i: [1,36,0,2,1,1,0,0] t3_sysrel
a: [3,68,0,2,1,36]
s: 68
i: [1,38,0,8,2,1,0,0] t3_sysrel
a: [3,54,0,4,1,38]
s: 54
i: [1,37,0,5,1,1,0,0] t3_sysrel
a: [3,63,0,5,1,37]
s: 63
i: [1,46,0,3,1,0,0,1] reflexivity_r1_tarski
s: 46
i: [1,45,0,4,1,0,0,255] existence_m1_subset_1
s: 45
i: [1,47,0,4,1,0,0,1] fc6_relat_1
s: 47
i: [1,48,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 48
i: [1,50,0,6,2,0,0,1] fc1_relat_1
s: 50
i: [1,41,0,7,2,0,0,255] t3_subset
s: 41
i: [1,42,0,7,2,0,0,255] t3_subset
s: 42
i: [1,49,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 49
d: [2,71,1,5,1,0,40,41,45]
s: 71
i: [1,39,0,8,3,0,0,255] cc2_relat_1
s: 39
d: [2,73,1,6,2,0,44,50,49]
s: 73
i: [1,40,0,8,2,0,0,255] t28_xboole_1
s: 40
i: [1,43,0,10,2,0,0,1] t1_sysrel
s: 43
d: [2,74,1,6,2,0,44,50,49]
s: 74
i: [1,44,0,10,2,0,0,1] t1_sysrel
s: 44
d: [2,75,1,6,2,0,40,39,45]
s: 75
d: [2,86,1,10,2,0,44,43,48]
d: [2,90,1,6,1,0,21,86,47]
s: 90
d: [2,91,1,10,2,0,44,44,48]
d: [2,95,1,6,1,0,21,91,47]
s: 95
d: [2,76,1,7,3,0,40,39,42]
s: 76
d: [2,80,1,8,2,0,44,40,49]
s: 80
d: [2,82,2,7,3,0,44,73,40]
s: 82
d: [2,81,1,8,2,0,44,40,49]
s: 81
d: [2,83,1,8,2,0,44,49,40]
s: 83
d: [2,98,2,7,2,0,40,76,90]
s: 98
d: [2,84,1,8,2,0,44,49,40]
s: 84
d: [2,100,2,7,2,0,40,76,95]
s: 100
d: [2,88,1,10,2,0,44,43,49]
s: 88
d: [2,89,1,10,2,0,44,43,49]
s: 89
d: [2,103,2,9,3,0,44,80,40]
s: 103
d: [2,93,1,10,2,0,44,44,49]
s: 93
d: [2,125,3,10,2,1,40,103,63]
a: [3,130,3,5,1,125]
s: 130
d: [2,94,1,10,2,0,44,44,49]
s: 94
d: [2,87,1,11,3,0,44,43,40]
s: 87
d: [2,143,2,0,0,3,52,68,54,63,87]
e: 143
d: [2,151,2,6,2,1,40,87,63]
d: [2,154,2,4,1,2,21,151,68]
d: [2,155,2,0,0,3,21,154,54]
e: 155
r: 54
a: [3,53,0,4,1,38]
s: 53
a: [3,58,0,4,1,38]
s: 58
d: [2,159,2,5,2,1,40,53,76]
a: [3,176,2,2,1,159]
s: 176
d: [2,92,1,11,3,0,44,44,40]
s: 92
d: [2,178,2,0,0,3,52,68,58,63,92]
e: 178
d: [2,186,2,6,2,1,40,92,63]
d: [2,189,2,4,1,2,21,186,68]
d: [2,190,2,0,0,3,21,189,58]
e: 190
f: 190,154,36,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for sysrel__t3_sysrel
% SZS output start Proof for sysrel__t3_sysrel
1. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X2,k2_zfmisc_1(X0,X1)) => (r1_tarski(k10_xtuple_0(X2),X1) & r1_tarski(k9_xtuple_0(X2),X0)))) [input t3_sysrel]
2. ~! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X2,k2_zfmisc_1(X0,X1)) => (r1_tarski(k10_xtuple_0(X2),X1) & r1_tarski(k9_xtuple_0(X2),X0)))) [negated conjecture 1]
16. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(k10_xtuple_0(k3_xboole_0(X2,k2_zfmisc_1(X0,X1))),X1) & r1_tarski(k9_xtuple_0(k3_xboole_0(X2,k2_zfmisc_1(X0,X1))),X0))) [input t1_sysrel]
17. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
19. ~! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(X2,k2_zfmisc_1(X0,X1)) => (r1_tarski(k10_xtuple_0(X2),X1) & r1_tarski(k9_xtuple_0(X2),X0)))) [flattening 2]
20. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 17]
22. ! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(k10_xtuple_0(k3_xboole_0(X2,k2_zfmisc_1(X0,X1))),X1) & r1_tarski(k9_xtuple_0(k3_xboole_0(X2,k2_zfmisc_1(X0,X1))),X0))) [flattening 16]
25. ? [X0,X1,X2] : (((~r1_tarski(k10_xtuple_0(X2),X1) | ~r1_tarski(k9_xtuple_0(X2),X0)) & r1_tarski(X2,k2_zfmisc_1(X0,X1))) & v1_relat_1(X2)) [ennf transformation 19]
26. ? [X0,X1,X2] : ((~r1_tarski(k10_xtuple_0(X2),X1) | ~r1_tarski(k9_xtuple_0(X2),X0)) & r1_tarski(X2,k2_zfmisc_1(X0,X1)) & v1_relat_1(X2)) [flattening 25]
28. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 20]
29. ! [X0,X1,X2] : ((r1_tarski(k10_xtuple_0(k3_xboole_0(X2,k2_zfmisc_1(X0,X1))),X1) & r1_tarski(k9_xtuple_0(k3_xboole_0(X2,k2_zfmisc_1(X0,X1))),X0)) | ~v1_relat_1(X2)) [ennf transformation 22]
31. ? [X0,X1,X2] : ((~r1_tarski(k10_xtuple_0(X2),X1) | ~r1_tarski(k9_xtuple_0(X2),X0)) & r1_tarski(X2,k2_zfmisc_1(X0,X1)) & v1_relat_1(X2)) => ((~r1_tarski(k10_xtuple_0(sK2),sK1) | ~r1_tarski(k9_xtuple_0(sK2),sK0)) & r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) & v1_relat_1(sK2)) [choice axiom]
32. (~r1_tarski(k10_xtuple_0(sK2),sK1) | ~r1_tarski(k9_xtuple_0(sK2),sK0)) & r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) & v1_relat_1(sK2) [skolemisation 26,31]
36. v1_relat_1(sK2) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) [cnf transformation 32] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. ~r1_tarski(k10_xtuple_0(sK2),sK1) | ~r1_tarski(k9_xtuple_0(sK2),sK0) [cnf transformation 32] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 28] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. r1_tarski(k9_xtuple_0(k3_xboole_0(X2,k2_zfmisc_1(X0,X1))),X0) | ~v1_relat_1(X2) [cnf transformation 29] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. r1_tarski(k10_xtuple_0(k3_xboole_0(X2,k2_zfmisc_1(X0,X1))),X1) | ~v1_relat_1(X2) [cnf transformation 29] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. 1 <=> r1_tarski(k9_xtuple_0(sK2),sK0) [avatar definition]
56. 2 <=> r1_tarski(k10_xtuple_0(sK2),sK1) [avatar definition]
58. ~r1_tarski(k10_xtuple_0(sK2),sK1) <- (~2) [avatar component clause 56] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
59. ~1 | ~2 [avatar split clause 38,56,52]
61. 3 <=> r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) [avatar definition]
63. r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) <- (3) [avatar component clause 61] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
64. 3 [avatar split clause 37,61]
66. 4 <=> v1_relat_1(sK2) [avatar definition]
68. v1_relat_1(sK2) <- (4) [avatar component clause 66] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
69. 4 [avatar split clause 36,66]
87. ~r1_tarski(X2,k2_zfmisc_1(X3,X4)) | ~v1_relat_1(X2) | r1_tarski(k9_xtuple_0(X2),X3) [superposition 43,40] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
92. ~r1_tarski(X2,k2_zfmisc_1(X3,X4)) | ~v1_relat_1(X2) | r1_tarski(k10_xtuple_0(X2),X4) [superposition 44,40] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
151. ~v1_relat_1(sK2) | r1_tarski(k9_xtuple_0(sK2),sK0) <- (3) [resolution 87,63] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
154. r1_tarski(k9_xtuple_0(sK2),sK0) <- (3, 4) [subsumption resolution 151,68] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
157. 1 | ~3 | ~4 [avatar split clause 154,66,61,52]
186. ~v1_relat_1(sK2) | r1_tarski(k10_xtuple_0(sK2),sK1) <- (3) [resolution 92,63] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
189. r1_tarski(k10_xtuple_0(sK2),sK1) <- (3, 4) [subsumption resolution 186,68] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
190. $false <- (~2, 3, 4) [subsumption resolution 189,58] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
191. 2 | ~3 | ~4 [avatar contradiction clause 190]
192. $false [avatar sat refutation 59,64,69,157,191] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:64,thDist:-64}
% SZS output end Proof for sysrel__t3_sysrel
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 121
% Active clauses: 33
% Passive clauses: 62
% Final active clauses: 32
% Final passive clauses: 22
% Main loop iterations started: 40
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Equational tautologies: 1
% Forward subsumptions: 28
% 
% >>> Generating Inferences
% Binary resolution: 42
% Unit resulting resolution: 2
% Forward superposition: 27
% Backward superposition: 13
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 2
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.001 s ( own 0.0000 s ) 

