i: [1,15,0,3,1,1,0,0] t10_xtuple_0
a: [3,28,0,3,1,15]
s: 28
i: [1,16,0,5,1,1,0,0] t10_xtuple_0
a: [3,23,0,5,1,16]
s: 23
i: [1,19,0,3,1,0,0,1] reflexivity_r1_tarski
s: 19
i: [1,17,0,6,1,0,0,255] d14_xtuple_0
s: 17
i: [1,18,0,8,2,0,0,255] t8_xtuple_0
s: 18
d: [2,30,1,7,1,0,44,17,17]
s: 30
d: [2,35,2,8,1,0,44,17,30]
s: 35
d: [2,31,1,9,2,0,44,18,17]
s: 31
d: [2,32,1,9,2,0,44,18,17]
s: 32
d: [2,33,2,10,2,0,44,18,30]
s: 33
d: [2,34,2,10,2,0,44,18,30]
s: 34
d: [2,36,3,9,1,0,44,35,17]
s: 36
d: [2,44,2,10,2,0,44,31,17]
s: 44
d: [2,74,3,8,2,0,40,44,18]
s: 74
d: [2,84,4,0,0,2,52,28,23,74]
e: 84
d: [2,85,4,3,1,1,40,74,23]
d: [2,87,4,0,0,2,21,85,28]
e: 87
f: 87,15,16
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t10_xtuple_0
% SZS output start Proof for xtuple_0__t10_xtuple_0
1. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1))) [input t10_xtuple_0]
2. ~! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1))) [negated conjecture 1]
3. ! [X0] : k11_xtuple_0(X0) = k9_xtuple_0(k9_xtuple_0(X0)) [input d14_xtuple_0]
7. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) [input t8_xtuple_0]
8. ~! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1))) [flattening 2]
9. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) [flattening 7]
11. ? [X0,X1] : (~r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1)) & r1_tarski(X0,X1)) [ennf transformation 8]
12. ! [X0,X1] : (r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 9]
13. ? [X0,X1] : (~r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1)) & r1_tarski(X0,X1)) => (~r1_tarski(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) & r1_tarski(sK0,sK1)) [choice axiom]
14. ~r1_tarski(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) & r1_tarski(sK0,sK1) [skolemisation 11,13]
15. r1_tarski(sK0,sK1) [cnf transformation 14] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
16. ~r1_tarski(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) [cnf transformation 14] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
17. k11_xtuple_0(X0) = k9_xtuple_0(k9_xtuple_0(X0)) [cnf transformation 3] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
18. r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 12] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
21. 1 <=> r1_tarski(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) [avatar definition]
23. ~r1_tarski(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) <- (~1) [avatar component clause 21] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
24. ~1 [avatar split clause 16,21]
26. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
28. r1_tarski(sK0,sK1) <- (2) [avatar component clause 26] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
29. 2 [avatar split clause 15,26]
31. r1_tarski(k11_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(k9_xtuple_0(X0),X1) [superposition 18,17] {a:1,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
44. ~r1_tarski(k9_xtuple_0(X1),k9_xtuple_0(X0)) | r1_tarski(k11_xtuple_0(X1),k11_xtuple_0(X0)) [superposition 31,17] {a:2,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
74. r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1)) | ~r1_tarski(X0,X1) [resolution 44,18] {a:3,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:4,thDist:-4}
85. ~r1_tarski(sK0,sK1) <- (~1) [resolution 74,23] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
87. $false <- (~1, 2) [subsumption resolution 85,28] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
88. 1 | ~2 [avatar contradiction clause 87]
89. $false [avatar sat refutation 24,29,88] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:20,thDist:-20}
% SZS output end Proof for xtuple_0__t10_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 6
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 5
% Generated clauses: 64
% Active clauses: 14
% Passive clauses: 54
% Final active clauses: 14
% Final passive clauses: 40
% Main loop iterations started: 14
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 1
% 
% >>> Deletion Inferences
% Forward subsumptions: 1
% Fw demodulations to eq. taut.: 4
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 1
% Forward superposition: 32
% Backward superposition: 18
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 6
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% passive container maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

