i: [1,46,0,3,1,1,0,0] t22_arytm_1
a: [3,68,0,3,1,46]
s: 68
i: [1,45,0,3,1,1,0,0] t22_arytm_1
a: [3,73,0,3,1,45]
s: 73
i: [1,44,0,3,1,1,0,0] t22_arytm_1
a: [3,78,0,3,1,44]
s: 78
i: [1,43,0,3,1,1,0,0] t22_arytm_1
a: [3,83,0,3,1,43]
s: 83
i: [1,42,0,3,1,1,0,0] t22_arytm_1
a: [3,88,0,3,1,42]
s: 88
i: [1,47,0,11,1,1,0,0] t22_arytm_1
a: [3,63,0,11,1,47]
s: 63
i: [1,48,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,93,0,3,1,48]
s: 93
i: [1,49,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,98,0,3,1,49]
s: 98
i: [1,55,0,4,1,0,0,1] existence_m1_subset_1
s: 55
i: [1,50,0,11,3,0,0,1] t11_arytm_1
s: 50
i: [1,56,0,11,3,0,0,1] dt_k7_arytm_2
s: 56
i: [1,57,0,11,3,0,0,1] dt_k1_arytm_1
s: 57
i: [1,59,0,12,4,0,0,1] connectedness_r1_arytm_2
s: 59
d: [2,102,1,9,3,1,40,59,88]
s: 102
i: [1,58,0,13,3,0,0,1] commutativity_k7_arytm_2
s: 58
i: [1,51,0,16,4,0,0,255] d2_arytm_1
s: 51
d: [2,108,2,6,2,1,40,102,88]
d: [2,112,2,3,1,1,18,108]
a: [3,116,2,3,1,112]
s: 116
i: [1,52,0,18,4,0,0,255] d2_arytm_1
s: 52
d: [2,110,2,6,2,2,40,102,78]
d: [2,118,2,3,1,3,21,110,68]
a: [3,122,2,3,1,118]
s: 122
i: [1,53,0,18,6,0,0,1] t3_arytm_1
s: 53
i: [1,54,0,26,6,0,0,1] l27_arytm_1
s: 54
d: [2,111,2,8,2,1,40,102,55]
a: [3,131,2,4,1,111]
s: 131
d: [2,133,1,19,3,1,44,63,58]
d: [2,139,1,16,2,2,21,133,78]
a: [3,143,1,5,1,139]
s: 143
d: [2,191,2,0,0,3,52,83,88,143,57]
e: 191
d: [2,192,2,6,2,1,40,143,57]
d: [2,194,2,3,1,2,21,192,88]
d: [2,195,2,0,0,3,21,194,83]
e: 195
r: 143
a: [3,142,1,5,1,139]
s: 142
d: [2,103,1,9,3,1,40,59,83]
s: 103
d: [2,104,1,9,3,1,40,59,78]
s: 104
d: [2,212,2,6,2,1,40,103,83]
d: [2,215,2,3,1,1,18,212]
a: [3,228,2,3,1,215]
s: 228
d: [2,151,1,15,5,1,40,53,73]
d: [2,157,1,12,4,2,21,151,83]
d: [2,158,1,9,3,3,21,157,88]
s: 158
d: [2,213,2,6,2,2,40,103,78]
a: [3,233,2,3,1,213]
s: 233
a: [3,147,1,11,1,139]
s: 147
d: [2,164,1,26,6,1,44,63,54]
d: [2,175,1,23,5,2,21,164,122]
d: [2,176,1,20,4,3,21,175,73]
d: [2,177,1,17,3,4,21,176,78]
d: [2,178,1,14,2,5,21,177,83]
d: [2,179,1,11,1,6,21,178,88]
a: [3,183,1,11,1,179]
s: 183
d: [2,253,2,6,2,1,40,104,78]
d: [2,255,2,3,1,1,18,253]
a: [3,269,2,3,1,255]
s: 269
d: [2,105,1,11,3,0,40,59,55]
s: 105
d: [2,214,2,8,2,1,40,103,55]
a: [3,242,2,4,1,214]
s: 242
d: [2,153,1,25,7,0,40,53,50]
d: [2,155,1,22,6,0,18,153]
d: [2,159,1,17,5,0,21,155,57]
s: 159
d: [2,100,1,19,5,0,40,59,57]
s: 100
d: [2,254,2,8,2,1,40,104,55]
a: [3,274,2,4,1,254]
s: 274
d: [2,336,3,0,0,5,52,88,68,78,131,55,274,53]
e: 336
r: 274
d: [2,101,1,19,5,0,40,59,56]
s: 101
a: [3,278,2,4,1,254]
s: 278
d: [2,167,1,30,8,0,44,56,54]
d: [2,168,1,27,7,0,18,167]
d: [2,187,1,22,6,0,21,168,57]
s: 187
d: [2,162,1,34,8,0,44,54,58]
d: [2,172,1,31,7,0,18,162]
d: [2,173,1,26,6,0,21,172,57]
s: 173
d: [2,197,2,10,2,2,40,142,102]
a: [3,202,2,5,1,197]
s: 202
d: [2,163,1,34,8,0,44,54,58]
d: [2,171,1,31,7,0,18,163]
d: [2,174,1,26,6,0,21,171,57]
s: 174
d: [2,209,2,10,2,2,40,103,142]
a: [3,219,2,5,1,209]
s: 219
d: [2,165,1,34,8,0,44,58,54]
d: [2,170,1,31,7,0,18,165]
d: [2,185,1,26,6,0,21,170,57]
s: 185
d: [2,166,1,34,8,0,44,58,54]
d: [2,169,1,31,7,0,18,166]
d: [2,186,1,26,6,0,21,169,57]
s: 186
d: [2,249,2,10,2,2,40,104,142]
a: [3,259,2,5,1,249]
s: 259
d: [2,291,2,24,4,1,44,183,51]
d: [2,292,2,13,3,1,19,291]
d: [2,293,2,10,2,2,21,292,83]
a: [3,297,2,5,1,293]
s: 297
d: [2,398,3,0,0,3,52,88,78,297,57]
e: 398
d: [2,399,3,6,2,1,40,297,57]
d: [2,401,3,3,1,2,21,399,78]
d: [2,402,3,0,0,3,21,401,88]
e: 402
r: 297
a: [3,296,2,5,1,293]
s: 296
a: [3,301,2,5,1,293]
s: 301
d: [2,198,2,13,3,1,40,142,59]
s: 198
d: [2,404,3,10,2,2,40,296,104]
a: [3,411,3,5,1,404]
s: 411
d: [2,450,4,0,0,6,52,233,78,83,301,296,411,53]
e: 450
r: 411
d: [2,282,2,16,4,3,40,158,50]
d: [2,285,2,13,3,3,21,282,88]
s: 285
d: [2,455,3,0,0,6,52,78,301,296,285]
e: 455
d: [2,456,3,8,2,4,40,285,296]
d: [2,461,3,3,1,5,21,456,301]
d: [2,462,3,0,0,6,21,461,78]
e: 462
f: 462,402,293,179,118,42,43,44,45,46,47
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t22_arytm_1
% SZS output start Proof for arytm_1__t22_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (r1_arytm_2(X0,X1) => (k2_arytm_1(X1,k1_arytm_1(X0,X2)) = k7_arytm_2(k1_arytm_1(X1,X0),X2) | r1_arytm_2(X0,X2)))))) [input t22_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (r1_arytm_2(X0,X1) => (k2_arytm_1(X1,k1_arytm_1(X0,X2)) = k7_arytm_2(k1_arytm_1(X1,X0),X2) | r1_arytm_2(X0,X2)))))) [negated conjecture 1]
4. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => (r1_arytm_2(X1,X0) | r1_arytm_2(X0,X1))) [input connectedness_r1_arytm_2]
5. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((~r1_arytm_2(X1,X0) => k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,k1_arytm_1(X1,X0))) & (r1_arytm_2(X1,X0) => k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1))))) [input d2_arytm_1]
7. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => m1_subset_1(k1_arytm_1(X0,X1),k2_arytm_2)) [input dt_k1_arytm_1]
16. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => ((r1_arytm_2(X2,X0) & r1_arytm_2(X0,X1)) => k7_arytm_2(k1_arytm_1(X1,X0),X2) = k1_arytm_1(X1,k1_arytm_1(X0,X2)))))) [input l27_arytm_1]
18. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => r1_arytm_2(k1_arytm_1(X0,X1),X0))) [input t11_arytm_1]
19. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => ((r1_arytm_2(X1,X2) & r1_arytm_2(X0,X1)) => r1_arytm_2(X0,X2))))) [input t3_arytm_1]
20. ? [X0] : (? [X1] : (? [X2] : (((k2_arytm_1(X1,k1_arytm_1(X0,X2)) != k7_arytm_2(k1_arytm_1(X1,X0),X2) & ~r1_arytm_2(X0,X2)) & r1_arytm_2(X0,X1)) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
21. ? [X0] : (? [X1] : (? [X2] : (k2_arytm_1(X1,k1_arytm_1(X0,X2)) != k7_arytm_2(k1_arytm_1(X1,X0),X2) & ~r1_arytm_2(X0,X2) & r1_arytm_2(X0,X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 20]
22. ! [X0] : (! [X1] : (r1_arytm_2(k1_arytm_1(X0,X1),X0) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 18]
23. ! [X0] : (! [X1] : (((k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,k1_arytm_1(X1,X0)) | r1_arytm_2(X1,X0)) & (k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 5]
24. ! [X0] : (! [X1] : (! [X2] : ((r1_arytm_2(X0,X2) | (~r1_arytm_2(X1,X2) | ~r1_arytm_2(X0,X1))) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 19]
25. ! [X0] : (! [X1] : (! [X2] : (r1_arytm_2(X0,X2) | ~r1_arytm_2(X1,X2) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 24]
26. ! [X0] : (! [X1] : (! [X2] : ((k7_arytm_2(k1_arytm_1(X1,X0),X2) = k1_arytm_1(X1,k1_arytm_1(X0,X2)) | (~r1_arytm_2(X2,X0) | ~r1_arytm_2(X0,X1))) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 16]
27. ! [X0] : (! [X1] : (! [X2] : (k7_arytm_2(k1_arytm_1(X1,X0),X2) = k1_arytm_1(X1,k1_arytm_1(X0,X2)) | ~r1_arytm_2(X2,X0) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 26]
30. ! [X0,X1] : (m1_subset_1(k1_arytm_1(X0,X1),k2_arytm_2) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 7]
31. ! [X0,X1] : (m1_subset_1(k1_arytm_1(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 30]
34. ! [X0,X1] : ((r1_arytm_2(X1,X0) | r1_arytm_2(X0,X1)) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 4]
35. ! [X0,X1] : (r1_arytm_2(X1,X0) | r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 34]
36. ? [X0] : (? [X1] : (? [X2] : (k2_arytm_1(X1,k1_arytm_1(X0,X2)) != k7_arytm_2(k1_arytm_1(X1,X0),X2) & ~r1_arytm_2(X0,X2) & r1_arytm_2(X0,X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (? [X2] : (k2_arytm_1(X1,k1_arytm_1(sK0,X2)) != k7_arytm_2(k1_arytm_1(X1,sK0),X2) & ~r1_arytm_2(sK0,X2) & r1_arytm_2(sK0,X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
37. ? [X1] : (? [X2] : (k2_arytm_1(X1,k1_arytm_1(sK0,X2)) != k7_arytm_2(k1_arytm_1(X1,sK0),X2) & ~r1_arytm_2(sK0,X2) & r1_arytm_2(sK0,X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) => (? [X2] : (k2_arytm_1(sK1,k1_arytm_1(sK0,X2)) != k7_arytm_2(k1_arytm_1(sK1,sK0),X2) & ~r1_arytm_2(sK0,X2) & r1_arytm_2(sK0,sK1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
38. ? [X2] : (k2_arytm_1(sK1,k1_arytm_1(sK0,X2)) != k7_arytm_2(k1_arytm_1(sK1,sK0),X2) & ~r1_arytm_2(sK0,X2) & r1_arytm_2(sK0,sK1) & m1_subset_1(X2,k2_arytm_2)) => (k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k7_arytm_2(k1_arytm_1(sK1,sK0),sK2) & ~r1_arytm_2(sK0,sK2) & r1_arytm_2(sK0,sK1) & m1_subset_1(sK2,k2_arytm_2)) [choice axiom]
39. ((k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k7_arytm_2(k1_arytm_1(sK1,sK0),sK2) & ~r1_arytm_2(sK0,sK2) & r1_arytm_2(sK0,sK1) & m1_subset_1(sK2,k2_arytm_2)) & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 21,38,37,36]
42. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
43. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
44. m1_subset_1(sK2,k2_arytm_2) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
45. r1_arytm_2(sK0,sK1) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
46. ~r1_arytm_2(sK0,sK2) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
47. k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k7_arytm_2(k1_arytm_1(sK1,sK0),sK2) [cnf transformation 39] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
50. r1_arytm_2(k1_arytm_1(X0,X1),X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 22] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1) | ~r1_arytm_2(X1,X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 23] {a:0,w:16,wCS:6400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
53. ~r1_arytm_2(X1,X2) | r1_arytm_2(X0,X2) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 25] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. k7_arytm_2(k1_arytm_1(X1,X0),X2) = k1_arytm_1(X1,k1_arytm_1(X0,X2)) | ~r1_arytm_2(X2,X0) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 27] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. m1_subset_1(k1_arytm_1(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 31] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. ~m1_subset_1(X1,k2_arytm_2) | r1_arytm_2(X0,X1) | r1_arytm_2(X1,X0) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 35] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. 1 <=> k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) = k7_arytm_2(k1_arytm_1(sK1,sK0),sK2) [avatar definition]
63. k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k7_arytm_2(k1_arytm_1(sK1,sK0),sK2) <- (~1) [avatar component clause 61] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
64. ~1 [avatar split clause 47,61]
66. 2 <=> r1_arytm_2(sK0,sK2) [avatar definition]
68. ~r1_arytm_2(sK0,sK2) <- (~2) [avatar component clause 66] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
69. ~2 [avatar split clause 46,66]
71. 3 <=> r1_arytm_2(sK0,sK1) [avatar definition]
73. r1_arytm_2(sK0,sK1) <- (3) [avatar component clause 71] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
74. 3 [avatar split clause 45,71]
76. 4 <=> m1_subset_1(sK2,k2_arytm_2) [avatar definition]
78. m1_subset_1(sK2,k2_arytm_2) <- (4) [avatar component clause 76] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
79. 4 [avatar split clause 44,76]
81. 5 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
83. m1_subset_1(sK1,k2_arytm_2) <- (5) [avatar component clause 81] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
84. 5 [avatar split clause 43,81]
86. 6 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
88. m1_subset_1(sK0,k2_arytm_2) <- (6) [avatar component clause 86] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
89. 6 [avatar split clause 42,86]
102. ~m1_subset_1(X6,k2_arytm_2) | r1_arytm_2(sK0,X6) | r1_arytm_2(X6,sK0) <- (6) [resolution 59,88] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
110. r1_arytm_2(sK0,sK2) | r1_arytm_2(sK2,sK0) <- (4, 6) [resolution 102,78] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
118. r1_arytm_2(sK2,sK0) <- (~2, 4, 6) [subsumption resolution 110,68] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
120. 10 <=> r1_arytm_2(sK2,sK0) [avatar definition]
122. r1_arytm_2(sK2,sK0) <- (10) [avatar component clause 120] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
123. 10 | 2 | ~4 | ~6 [avatar split clause 118,86,76,66,120]
151. r1_arytm_2(X0,sK1) | ~r1_arytm_2(X0,sK0) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) <- (3) [resolution 53,73] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
157. r1_arytm_2(X0,sK1) | ~r1_arytm_2(X0,sK0) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) <- (3, 5) [subsumption resolution 151,83] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
158. ~r1_arytm_2(X0,sK0) | r1_arytm_2(X0,sK1) | ~m1_subset_1(X0,k2_arytm_2) <- (3, 5, 6) [subsumption resolution 157,88] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
164. k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) | ~r1_arytm_2(sK2,sK0) | ~r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1) [superposition 63,54] {a:1,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
175. k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) | ~r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1, 10) [subsumption resolution 164,122] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
176. k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1, 3, 10) [subsumption resolution 175,73] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
177. k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1, 3, 4, 10) [subsumption resolution 176,78] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
178. k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1, 3, 4, 5, 10) [subsumption resolution 177,83] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
179. k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) <- (~1, 3, 4, 5, 6, 10) [subsumption resolution 178,88] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
181. 15 <=> k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) = k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) [avatar definition]
183. k2_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) <- (~15) [avatar component clause 181] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
184. ~15 | 1 | ~3 | ~4 | ~5 | ~6 | ~10 [avatar split clause 179,120,86,81,76,71,61,181]
282. r1_arytm_2(k1_arytm_1(sK0,X0),sK1) | ~m1_subset_1(k1_arytm_1(sK0,X0),k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (3, 5, 6) [resolution 158,50] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
285. ~m1_subset_1(k1_arytm_1(sK0,X0),k2_arytm_2) | r1_arytm_2(k1_arytm_1(sK0,X0),sK1) | ~m1_subset_1(X0,k2_arytm_2) <- (3, 5, 6) [subsumption resolution 282,88] {a:2,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
291. k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) != k1_arytm_1(sK1,k1_arytm_1(sK0,sK2)) | ~r1_arytm_2(k1_arytm_1(sK0,sK2),sK1) | ~m1_subset_1(k1_arytm_1(sK0,sK2),k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (~15) [superposition 183,51] {a:2,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
292. ~r1_arytm_2(k1_arytm_1(sK0,sK2),sK1) | ~m1_subset_1(k1_arytm_1(sK0,sK2),k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (~15) [trivial inequality removal 291] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
293. ~r1_arytm_2(k1_arytm_1(sK0,sK2),sK1) | ~m1_subset_1(k1_arytm_1(sK0,sK2),k2_arytm_2) <- (5, ~15) [subsumption resolution 292,83] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
295. 30 <=> m1_subset_1(k1_arytm_1(sK0,sK2),k2_arytm_2) [avatar definition]
296. m1_subset_1(k1_arytm_1(sK0,sK2),k2_arytm_2) <- (30) [avatar component clause 295] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
297. ~m1_subset_1(k1_arytm_1(sK0,sK2),k2_arytm_2) <- (~30) [avatar component clause 295] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
299. 31 <=> r1_arytm_2(k1_arytm_1(sK0,sK2),sK1) [avatar definition]
301. ~r1_arytm_2(k1_arytm_1(sK0,sK2),sK1) <- (~31) [avatar component clause 299] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
302. ~30 | ~31 | ~5 | 15 [avatar split clause 293,181,81,299,295]
399. ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~30) [resolution 297,57] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
401. ~m1_subset_1(sK0,k2_arytm_2) <- (4, ~30) [subsumption resolution 399,78] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
402. $false <- (4, 6, ~30) [subsumption resolution 401,88] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
403. ~4 | ~6 | 30 [avatar contradiction clause 402]
456. r1_arytm_2(k1_arytm_1(sK0,sK2),sK1) | ~m1_subset_1(sK2,k2_arytm_2) <- (3, 5, 6, 30) [resolution 285,296] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
461. ~m1_subset_1(sK2,k2_arytm_2) <- (3, 5, 6, 30, ~31) [subsumption resolution 456,301] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
462. $false <- (3, 4, 5, 6, 30, ~31) [subsumption resolution 461,78] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
463. ~3 | ~4 | ~5 | ~6 | ~30 | 31 [avatar contradiction clause 462]
467. $false [avatar sat refutation 64,69,74,79,84,89,123,184,302,403,463] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:475,thDist:-475}
% SZS output end Proof for arytm_1__t22_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 18
% Generated clauses: 267
% Active clauses: 50
% Passive clauses: 95
% Final active clauses: 46
% Final passive clauses: 42
% Main loop iterations started: 53
% 
% >>> Simplifying Inferences
% Duplicate literals: 32
% Trivial inequalities: 1
% Fw subsumption resolutions: 80
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 19
% 
% >>> Generating Inferences
% Binary resolution: 74
% Unit resulting resolution: 5
% Factoring: 4
% Forward superposition: 9
% Backward superposition: 14
% 
% >>> AVATAR
% Split clauses: 14
% Split components: 28
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 42
% SAT solver unit clauses: 9
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% superposition: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

