i: [1,48,0,15,3,1,0,0] t63_zfmisc_1
a: [3,73,0,3,1,48]
s: 73
a: [3,77,0,3,1,48]
s: 77
i: [1,49,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,85,0,2,1,49]
s: 85
i: [1,50,0,2,1,0,0,2] fc1_xboole_0
a: [3,90,0,2,1,50]
s: 90
i: [1,65,0,2,1,0,0,2] rc2_xboole_0
a: [3,100,0,2,1,65]
s: 100
i: [1,66,0,2,1,0,0,2] rc1_xboole_0
a: [3,105,0,2,1,66]
s: 105
a: [3,70,0,9,1,48]
s: 70
i: [1,51,0,3,1,0,0,255] d2_xboole_0
a: [3,95,0,3,1,51]
s: 95
r: 90
i: [1,61,0,4,1,0,0,1] fc3_xboole_0
s: 61
i: [1,52,0,5,1,0,0,1] t4_boole
s: 52
i: [1,53,0,5,1,0,0,1] t3_boole
s: 53
i: [1,54,0,5,2,0,0,2] t6_boole
s: 54
i: [1,58,0,5,2,0,0,1] t7_boole
s: 58
d: [2,108,0,5,1,1,22,52,95]
s: 108
d: [2,109,0,5,1,1,22,53,95]
s: 109
d: [2,110,0,5,2,1,22,54,95]
s: 110
i: [1,63,0,6,2,0,0,2] symmetry_r1_xboole_0
s: 63
i: [1,64,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 64
i: [1,57,0,7,3,0,0,1] t8_boole
s: 57
i: [1,62,0,7,1,0,0,1] commutativity_k2_tarski
s: 62
i: [1,55,0,8,2,0,0,1] t83_xboole_1
s: 55
d: [2,117,1,12,2,1,44,70,55]
d: [2,118,1,5,1,1,19,117]
a: [3,122,1,5,1,118]
s: 122
i: [1,56,0,8,2,0,0,1] t83_xboole_1
s: 56
d: [2,125,1,6,2,1,44,56,109]
d: [2,127,1,3,1,1,19,125]
s: 127
i: [1,60,0,8,2,0,0,2] t49_zfmisc_1
s: 60
i: [1,59,0,11,3,0,0,2] t51_zfmisc_1
s: 59
d: [2,134,2,0,0,3,52,73,77,122,59]
e: 134
d: [2,135,2,6,2,1,40,59,122]
d: [2,141,2,3,1,2,21,135,77]
d: [2,142,2,0,0,3,21,141,73]
e: 142
r: 70
r: 73
r: 122
a: [3,74,0,3,1,48]
s: 74
a: [3,69,0,9,1,48]
s: 69
d: [2,145,1,2,1,1,40,74,58]
a: [3,154,1,2,1,145]
s: 154
d: [2,144,1,3,1,1,40,74,64]
a: [3,149,1,3,1,144]
s: 149
a: [3,121,1,5,1,118]
s: 121
d: [2,159,2,0,0,2,52,74,121,60]
e: 159
d: [2,160,2,3,1,1,40,121,60]
d: [2,163,2,0,0,2,21,160,74]
e: 163
r: 74
r: 77
s: 73
a: [3,78,0,3,1,48]
s: 78
d: [2,171,1,3,1,1,40,78,64]
a: [3,176,1,3,1,171]
s: 176
d: [2,124,1,6,2,1,44,56,108]
d: [2,128,1,3,1,1,19,124]
s: 128
d: [2,111,1,7,3,1,44,110,110]
s: 111
d: [2,161,2,5,1,1,40,121,63]
a: [3,169,2,5,1,161]
s: 169
d: [2,112,1,7,2,1,44,109,110]
s: 112
d: [2,129,2,3,1,1,40,127,63]
s: 129
d: [2,113,1,7,2,1,44,108,110]
s: 113
d: [2,132,1,8,2,0,44,60,62]
s: 132
d: [2,185,2,0,0,2,52,78,121,132]
e: 185
d: [2,187,2,3,1,1,40,132,121]
d: [2,192,2,0,0,2,21,187,78]
e: 192
i: [1,47,0,12,2,1,0,0] t63_zfmisc_1
i: [1,46,0,12,2,1,0,0] t63_zfmisc_1
d: [2,156,1,12,2,1,44,56,69]
d: [2,157,1,5,1,1,19,156]
f: 192,160,157,142,118,46,47,48
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t63_zfmisc_1
% SZS output start Proof for zfmisc_1__t63_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = k4_xboole_0(k2_tarski(X0,X1),X2) <=> (~r2_hidden(X1,X2) & ~r2_hidden(X0,X2))) [input t63_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = k4_xboole_0(k2_tarski(X0,X1),X2) <=> (~r2_hidden(X1,X2) & ~r2_hidden(X0,X2))) [negated conjecture 1]
13. ! [X0,X1] : k2_tarski(X0,X1) = k2_tarski(X1,X0) [input commutativity_k2_tarski]
20. ! [X0] : ! [X1] : ! [X2] : ~(r2_hidden(X0,X2) & r1_xboole_0(k2_tarski(X0,X1),X2)) [input t49_zfmisc_1]
21. ! [X0] : ! [X1] : ! [X2] : ~(~r1_xboole_0(k2_tarski(X0,X2),X1) & ~r2_hidden(X2,X1) & ~r2_hidden(X0,X1)) [input t51_zfmisc_1]
22. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [input t83_xboole_1]
23. ~! [X0,X1,X2] : (k2_tarski(X0,X1) = k4_xboole_0(k2_tarski(X0,X1),X2) <=> (~r2_hidden(X1,X2) & ~r2_hidden(X0,X2))) [flattening 2]
24. ! [X0,X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [flattening 22]
27. ! [X0,X1,X2] : ~(~r1_xboole_0(k2_tarski(X0,X2),X1) & ~r2_hidden(X2,X1) & ~r2_hidden(X0,X1)) [flattening 21]
28. ! [X0,X1,X2] : ~(r2_hidden(X0,X2) & r1_xboole_0(k2_tarski(X0,X1),X2)) [flattening 20]
29. ? [X0,X1,X2] : (k2_tarski(X0,X1) = k4_xboole_0(k2_tarski(X0,X1),X2) <~> (~r2_hidden(X1,X2) & ~r2_hidden(X0,X2))) [ennf transformation 23]
33. ! [X0,X1,X2] : (r1_xboole_0(k2_tarski(X0,X2),X1) | r2_hidden(X2,X1) | r2_hidden(X0,X1)) [ennf transformation 27]
34. ! [X0,X1,X2] : (~r2_hidden(X0,X2) | ~r1_xboole_0(k2_tarski(X0,X1),X2)) [ennf transformation 28]
37. ? [X0,X1,X2] : (((r2_hidden(X1,X2) | r2_hidden(X0,X2)) | k2_tarski(X0,X1) != k4_xboole_0(k2_tarski(X0,X1),X2)) & ((~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) | k2_tarski(X0,X1) = k4_xboole_0(k2_tarski(X0,X1),X2))) [nnf transformation 29]
38. ? [X0,X1,X2] : ((r2_hidden(X1,X2) | r2_hidden(X0,X2) | k2_tarski(X0,X1) != k4_xboole_0(k2_tarski(X0,X1),X2)) & ((~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) | k2_tarski(X0,X1) = k4_xboole_0(k2_tarski(X0,X1),X2))) [flattening 37]
39. ? [X0,X1,X2] : ((r2_hidden(X1,X2) | r2_hidden(X0,X2) | k2_tarski(X0,X1) != k4_xboole_0(k2_tarski(X0,X1),X2)) & ((~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) | k2_tarski(X0,X1) = k4_xboole_0(k2_tarski(X0,X1),X2))) => ((r2_hidden(sK1,sK2) | r2_hidden(sK0,sK2) | k2_tarski(sK0,sK1) != k4_xboole_0(k2_tarski(sK0,sK1),sK2)) & ((~r2_hidden(sK1,sK2) & ~r2_hidden(sK0,sK2)) | k2_tarski(sK0,sK1) = k4_xboole_0(k2_tarski(sK0,sK1),sK2))) [choice axiom]
40. (r2_hidden(sK1,sK2) | r2_hidden(sK0,sK2) | k2_tarski(sK0,sK1) != k4_xboole_0(k2_tarski(sK0,sK1),sK2)) & ((~r2_hidden(sK1,sK2) & ~r2_hidden(sK0,sK2)) | k2_tarski(sK0,sK1) = k4_xboole_0(k2_tarski(sK0,sK1),sK2)) [skolemisation 38,39]
41. ! [X0,X1] : ((r1_xboole_0(X0,X1) | k4_xboole_0(X0,X1) != X0) & (k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1))) [nnf transformation 24]
46. ~r2_hidden(sK0,sK2) | k2_tarski(sK0,sK1) = k4_xboole_0(k2_tarski(sK0,sK1),sK2) [cnf transformation 40] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. ~r2_hidden(sK1,sK2) | k2_tarski(sK0,sK1) = k4_xboole_0(k2_tarski(sK0,sK1),sK2) [cnf transformation 40] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. r2_hidden(sK1,sK2) | r2_hidden(sK0,sK2) | k2_tarski(sK0,sK1) != k4_xboole_0(k2_tarski(sK0,sK1),sK2) [cnf transformation 40] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1) [cnf transformation 41] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. k4_xboole_0(X0,X1) != X0 | r1_xboole_0(X0,X1) [cnf transformation 41] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. r1_xboole_0(k2_tarski(X0,X2),X1) | r2_hidden(X2,X1) | r2_hidden(X0,X1) [cnf transformation 33] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
60. ~r1_xboole_0(k2_tarski(X0,X1),X2) | ~r2_hidden(X0,X2) [cnf transformation 34] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
62. k2_tarski(X0,X1) = k2_tarski(X1,X0) [cnf transformation 13] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. 1 <=> k2_tarski(sK0,sK1) = k4_xboole_0(k2_tarski(sK0,sK1),sK2) [avatar definition]
69. k2_tarski(sK0,sK1) = k4_xboole_0(k2_tarski(sK0,sK1),sK2) <- (1) [avatar component clause 68] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
70. k2_tarski(sK0,sK1) != k4_xboole_0(k2_tarski(sK0,sK1),sK2) <- (~1) [avatar component clause 68] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. 2 <=> r2_hidden(sK0,sK2) [avatar definition]
73. ~r2_hidden(sK0,sK2) <- (~2) [avatar component clause 72] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. 3 <=> r2_hidden(sK1,sK2) [avatar definition]
77. ~r2_hidden(sK1,sK2) <- (~3) [avatar component clause 76] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
78. r2_hidden(sK1,sK2) <- (3) [avatar component clause 76] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. ~1 | 2 | 3 [avatar split clause 48,76,72,68]
80. 1 | ~3 [avatar split clause 47,76,68]
81. 1 | ~2 [avatar split clause 46,72,68]
117. k2_tarski(sK0,sK1) != k2_tarski(sK0,sK1) | ~r1_xboole_0(k2_tarski(sK0,sK1),sK2) <- (~1) [superposition 70,55] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
118. ~r1_xboole_0(k2_tarski(sK0,sK1),sK2) <- (~1) [trivial inequality removal 117] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
120. 9 <=> r1_xboole_0(k2_tarski(sK0,sK1),sK2) [avatar definition]
121. r1_xboole_0(k2_tarski(sK0,sK1),sK2) <- (9) [avatar component clause 120] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
122. ~r1_xboole_0(k2_tarski(sK0,sK1),sK2) <- (~9) [avatar component clause 120] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
123. ~9 | 1 [avatar split clause 118,68,120]
132. ~r1_xboole_0(k2_tarski(X1,X0),X2) | ~r2_hidden(X0,X2) [superposition 60,62] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
135. r2_hidden(sK1,sK2) | r2_hidden(sK0,sK2) <- (~9) [resolution 59,122] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
141. r2_hidden(sK0,sK2) <- (~3, ~9) [subsumption resolution 135,77] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
142. $false <- (~2, ~3, ~9) [subsumption resolution 141,73] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
143. 2 | 3 | 9 [avatar contradiction clause 142]
156. k2_tarski(sK0,sK1) != k2_tarski(sK0,sK1) | r1_xboole_0(k2_tarski(sK0,sK1),sK2) <- (1) [superposition 56,69] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
157. r1_xboole_0(k2_tarski(sK0,sK1),sK2) <- (1) [trivial inequality removal 156] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
158. 9 | ~1 [avatar split clause 157,68,120]
160. ~r2_hidden(sK0,sK2) <- (9) [resolution 121,60] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
165. ~2 | ~9 [avatar split clause 160,120,72]
187. ~r2_hidden(sK1,sK2) <- (9) [resolution 132,121] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
192. $false <- (3, 9) [subsumption resolution 187,78] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
193. ~3 | ~9 [avatar contradiction clause 192]
194. $false [avatar sat refutation 79,80,81,123,143,158,165,193] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:96,thDist:-96}
% SZS output end Proof for zfmisc_1__t63_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 96
% Active clauses: 36
% Passive clauses: 48
% Final active clauses: 30
% Final passive clauses: 7
% Main loop iterations started: 41
% 
% >>> Simplifying Inferences
% Trivial inequalities: 6
% Fw subsumption resolutions: 4
% Fw demodulations: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 16
% Unit resulting resolution: 3
% Forward superposition: 11
% Backward superposition: 9
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 5
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

