i: [1,111,0,2,1,1,0,0] t36_diraf
a: [3,156,0,2,1,111]
s: 156
i: [1,110,0,2,1,1,0,0] t36_diraf
a: [3,161,0,2,1,110]
s: 161
i: [1,109,0,2,1,1,0,0] t36_diraf
a: [3,166,0,2,1,109]
s: 166
i: [1,113,0,2,1,0,0,255] spc1_boole
a: [3,171,0,2,1,113]
s: 171
i: [1,114,0,2,1,0,0,255] fc1_xboole_0
a: [3,176,0,2,1,114]
s: 176
i: [1,149,0,2,1,0,0,2] rc2_xboole_0
a: [3,181,0,2,1,149]
s: 181
i: [1,150,0,2,1,0,0,1] existence_l1_analoaf
a: [3,186,0,2,1,150]
s: 186
i: [1,151,0,2,1,0,0,2] existence_l1_struct_0
a: [3,191,0,2,1,151]
s: 191
i: [1,152,0,2,1,0,0,2] rc1_xboole_0
a: [3,196,0,2,1,152]
s: 196
i: [1,115,0,4,2,0,0,1] dt_l1_analoaf
s: 115
i: [1,147,0,4,1,0,0,1] existence_m1_subset_1
s: 147
d: [2,198,1,2,1,1,40,115,156]
a: [3,203,1,2,1,198]
s: 203
i: [1,112,0,17,4,1,0,0] t36_diraf
s: 112
d: [2,199,1,2,1,1,40,115,186]
a: [3,208,1,2,1,199]
s: 208
i: [1,130,0,5,2,0,0,255] t6_boole
s: 130
i: [1,146,0,5,2,0,0,255] t7_boole
s: 146
i: [1,116,0,6,3,0,0,255] cc7_struct_0
s: 116
i: [1,117,0,6,3,0,0,255] cc5_struct_0
s: 117
d: [2,212,1,4,2,1,40,116,191]
a: [3,218,1,2,1,212]
s: 218
i: [1,119,0,6,3,0,0,1] cc17_struct_0
s: 119
i: [1,120,0,6,3,0,0,1] cc2_struct_0
s: 120
d: [2,214,2,4,2,1,40,116,208]
a: [3,227,2,2,1,214]
s: 227
i: [1,121,0,6,3,0,0,1] cc1_struct_0
s: 121
i: [1,124,0,6,3,0,0,255] cc4_struct_0
s: 124
i: [1,125,0,6,3,0,0,255] cc6_struct_0
s: 125
i: [1,143,0,6,2,0,0,255] t1_subset
s: 143
i: [1,148,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 148
i: [1,122,0,7,3,0,0,255] cc8_struct_0
s: 122
i: [1,126,0,7,3,0,0,255] cc11_struct_0
s: 126
i: [1,127,0,7,3,0,0,255] cc11_struct_0
s: 127
i: [1,128,0,7,3,0,0,255] cc9_struct_0
s: 128
i: [1,131,0,7,3,0,0,255] fc9_struct_0
s: 131
i: [1,134,0,7,3,0,0,1] fc2_struct_0
s: 134
i: [1,139,0,7,3,0,0,255] fc6_struct_0
s: 139
i: [1,140,0,7,3,0,0,255] fc8_struct_0
s: 140
i: [1,141,0,7,3,0,0,1] fc1_struct_0
s: 141
i: [1,142,0,7,3,0,0,255] fc7_struct_0
s: 142
i: [1,145,0,7,3,0,0,255] t8_boole
s: 145
d: [2,210,1,7,3,0,44,130,130]
s: 210
i: [1,144,0,8,3,0,0,255] t2_subset
s: 144
i: [1,129,0,9,4,0,0,255] cc10_struct_0
s: 129
d: [2,243,1,6,2,0,40,144,147]
s: 243
i: [1,135,0,11,4,0,0,1] t24_diraf
s: 135
d: [2,248,2,6,2,0,40,243,148]
s: 248
i: [1,136,0,11,4,0,0,1] t24_diraf
s: 136
i: [1,137,0,11,4,0,0,1] t24_diraf
s: 137
d: [2,233,1,9,4,0,44,122,130]
s: 233
i: [1,138,0,16,4,0,0,1] t24_diraf
s: 138
d: [2,235,1,9,4,0,44,128,130]
s: 235
i: [1,132,0,27,7,0,0,1] d5_diraf
s: 132
i: [1,133,0,27,7,0,0,1] d5_diraf
s: 133
d: [2,251,1,14,5,0,40,135,144]
s: 251
d: [2,252,1,14,5,0,40,136,144]
s: 252
d: [2,253,1,14,5,0,40,137,144]
s: 253
d: [2,266,1,33,9,0,40,132,138]
d: [2,267,1,31,8,0,18,266]
d: [2,268,1,26,7,0,21,267,137]
d: [2,269,1,21,6,0,21,268,136]
d: [2,270,1,16,5,0,21,269,135]
s: 270
d: [2,282,2,23,7,0,40,270,112]
d: [2,283,2,21,6,1,21,282,156]
d: [2,284,2,19,5,2,21,283,161]
d: [2,285,2,17,4,3,21,284,166]
a: [3,289,2,5,1,285]
s: 289
d: [2,303,3,0,0,4,52,166,161,156,289,135]
e: 303
d: [2,304,3,6,3,1,40,289,135]
d: [2,306,3,4,2,2,21,304,156]
d: [2,307,3,2,1,3,21,306,161]
d: [2,308,3,0,0,4,21,307,166]
e: 308
r: 289
a: [3,288,2,5,1,285]
s: 288
a: [3,293,2,5,1,285]
s: 293
d: [2,320,3,0,0,4,52,166,161,156,293,136]
e: 320
d: [2,321,3,6,3,1,40,293,136]
d: [2,323,3,4,2,2,21,321,156]
d: [2,324,3,2,1,3,21,323,161]
d: [2,325,3,0,0,4,21,324,166]
e: 325
r: 293
a: [3,292,2,5,1,285]
s: 292
a: [3,297,2,5,1,285]
s: 297
d: [2,333,3,0,0,4,52,166,161,156,297,137]
e: 333
d: [2,334,3,6,3,1,40,297,137]
d: [2,336,3,4,2,2,21,334,156]
d: [2,337,3,2,1,3,21,336,161]
d: [2,338,3,0,0,4,21,337,166]
e: 338
r: 297
a: [3,301,2,2,1,285]
s: 301
d: [2,340,3,0,0,3,52,203,166,301,119]
e: 340
d: [2,341,3,4,2,1,40,301,119]
d: [2,344,3,2,1,2,21,341,166]
d: [2,345,3,0,0,3,21,344,203]
e: 345
f: 345,338,325,308,285,198,109,110,111
% Refutation found. Thanks to Tanya!
% SZS status Theorem for diraf__t36_diraf
% SZS output start Proof for diraf__t36_diraf
1. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ~! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => r3_diraf(X0,X1,X2,X3))))) [input t36_diraf]
2. ~! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ~! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => r3_diraf(X0,X1,X2,X3))))) [negated conjecture 1]
6. ! [X0] : (l1_struct_0(X0) => (~v7_struct_0(X0) => ~v2_struct_0(X0))) [input cc17_struct_0]
15. ! [X0] : ((l1_analoaf(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r3_diraf(X0,X1,X2,X3) <=> r2_diraf(X0,X1,X2,X1,X3)))))) [input d5_diraf]
17. ! [X0] : (l1_analoaf(X0) => l1_struct_0(X0)) [input dt_l1_analoaf]
35. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ~! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => r2_diraf(X0,X1,X2,X1,X3))))) [input t24_diraf]
44. ? [X0] : (! [X1] : (! [X2] : (! [X3] : (r3_diraf(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
45. ? [X0] : (! [X1] : (! [X2] : (! [X3] : (r3_diraf(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) [flattening 44]
46. ! [X0] : (l1_struct_0(X0) | ~l1_analoaf(X0)) [ennf transformation 17]
51. ! [X0] : ((~v2_struct_0(X0) | v7_struct_0(X0)) | ~l1_struct_0(X0)) [ennf transformation 6]
52. ! [X0] : (~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0)) [flattening 51]
72. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r3_diraf(X0,X1,X2,X3) <=> r2_diraf(X0,X1,X2,X1,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | v2_struct_0(X0))) [ennf transformation 15]
73. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r3_diraf(X0,X1,X2,X3) <=> r2_diraf(X0,X1,X2,X1,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 72]
76. ! [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_diraf(X0,X1,X2,X1,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0))) [ennf transformation 35]
77. ! [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_diraf(X0,X1,X2,X1,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0)) [flattening 76]
92. ? [X0] : (! [X1] : (! [X2] : (! [X3] : (r3_diraf(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => (! [X1] : (! [X2] : (! [X3] : (r3_diraf(sK0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(sK0))) | ~m1_subset_1(X2,u1_struct_0(sK0))) | ~m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
93. ! [X1] : (! [X2] : (! [X3] : (r3_diraf(sK0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(sK0))) | ~m1_subset_1(X2,u1_struct_0(sK0))) | ~m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0) [skolemisation 45,92]
94. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r3_diraf(X0,X1,X2,X3) | ~r2_diraf(X0,X1,X2,X1,X3)) & (r2_diraf(X0,X1,X2,X1,X3) | ~r3_diraf(X0,X1,X2,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [nnf transformation 73]
95. ! [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_diraf(X0,X1,X2,X1,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) => (? [X2] : (? [X3] : (~r2_diraf(X0,sK1(X0),X2,sK1(X0),X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(sK1(X0),u1_struct_0(X0)))) [choice axiom]
96. ! [X0] : (? [X2] : (? [X3] : (~r2_diraf(X0,sK1(X0),X2,sK1(X0),X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) => (? [X3] : (~r2_diraf(X0,sK1(X0),sK2(X0),sK1(X0),X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(sK2(X0),u1_struct_0(X0)))) [choice axiom]
97. ! [X0] : (? [X3] : (~r2_diraf(X0,sK1(X0),sK2(X0),sK1(X0),X3) & m1_subset_1(X3,u1_struct_0(X0))) => (~r2_diraf(X0,sK1(X0),sK2(X0),sK1(X0),sK3(X0)) & m1_subset_1(sK3(X0),u1_struct_0(X0)))) [choice axiom]
98. ! [X0] : ((((~r2_diraf(X0,sK1(X0),sK2(X0),sK1(X0),sK3(X0)) & m1_subset_1(sK3(X0),u1_struct_0(X0))) & m1_subset_1(sK2(X0),u1_struct_0(X0))) & m1_subset_1(sK1(X0),u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0)) [skolemisation 77,97,96,95]
109. ~v7_struct_0(sK0) [cnf transformation 93] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
110. v2_analoaf(sK0) [cnf transformation 93] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
111. l1_analoaf(sK0) [cnf transformation 93] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
112. r3_diraf(sK0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(sK0)) | ~m1_subset_1(X2,u1_struct_0(sK0)) | ~m1_subset_1(X1,u1_struct_0(sK0)) [cnf transformation 93] {a:0,w:17,wCS:1700,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
115. ~l1_analoaf(X0) | l1_struct_0(X0) [cnf transformation 46] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
119. ~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 52] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
132. r2_diraf(X0,X1,X2,X1,X3) | ~r3_diraf(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 94] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
135. m1_subset_1(sK1(X0),u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 98] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
136. m1_subset_1(sK2(X0),u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 98] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
137. m1_subset_1(sK3(X0),u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 98] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
138. ~r2_diraf(X0,sK1(X0),sK2(X0),sK1(X0),sK3(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 98] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
154. 1 <=> l1_analoaf(sK0) [avatar definition]
156. l1_analoaf(sK0) <- (1) [avatar component clause 154] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
157. 1 [avatar split clause 111,154]
159. 2 <=> v2_analoaf(sK0) [avatar definition]
161. v2_analoaf(sK0) <- (2) [avatar component clause 159] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
162. 2 [avatar split clause 110,159]
164. 3 <=> v7_struct_0(sK0) [avatar definition]
166. ~v7_struct_0(sK0) <- (~3) [avatar component clause 164] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
167. ~3 [avatar split clause 109,164]
198. l1_struct_0(sK0) <- (1) [resolution 115,156] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
201. 10 <=> l1_struct_0(sK0) [avatar definition]
203. l1_struct_0(sK0) <- (10) [avatar component clause 201] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
204. 10 | ~1 [avatar split clause 198,154,201]
266. ~r3_diraf(X0,sK1(X0),sK2(X0),sK3(X0)) | ~m1_subset_1(sK3(X0),u1_struct_0(X0)) | ~m1_subset_1(sK2(X0),u1_struct_0(X0)) | ~m1_subset_1(sK1(X0),u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [resolution 132,138] {a:1,w:33,wCS:13200,sine:1,thAx:0,allAx:5,thDist:-5}
267. ~r3_diraf(X0,sK1(X0),sK2(X0),sK3(X0)) | ~m1_subset_1(sK3(X0),u1_struct_0(X0)) | ~m1_subset_1(sK2(X0),u1_struct_0(X0)) | ~m1_subset_1(sK1(X0),u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [duplicate literal removal 266] {a:1,w:31,wCS:12400,sine:1,thAx:0,allAx:5,thDist:-5}
268. ~r3_diraf(X0,sK1(X0),sK2(X0),sK3(X0)) | ~m1_subset_1(sK2(X0),u1_struct_0(X0)) | ~m1_subset_1(sK1(X0),u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [subsumption resolution 267,137] {a:1,w:26,wCS:10400,sine:1,thAx:0,allAx:5,thDist:-5}
269. ~r3_diraf(X0,sK1(X0),sK2(X0),sK3(X0)) | ~m1_subset_1(sK1(X0),u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [subsumption resolution 268,136] {a:1,w:21,wCS:8400,sine:1,thAx:0,allAx:5,thDist:-5}
270. ~r3_diraf(X0,sK1(X0),sK2(X0),sK3(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [subsumption resolution 269,135] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
282. ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK1(sK0),u1_struct_0(sK0)) [resolution 270,112] {a:2,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
283. v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK1(sK0),u1_struct_0(sK0)) <- (1) [subsumption resolution 282,156] {a:2,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
284. v2_struct_0(sK0) | v7_struct_0(sK0) | ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK1(sK0),u1_struct_0(sK0)) <- (1, 2) [subsumption resolution 283,161] {a:2,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
285. v2_struct_0(sK0) | ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK1(sK0),u1_struct_0(sK0)) <- (1, 2, ~3) [subsumption resolution 284,166] {a:2,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
287. 16 <=> m1_subset_1(sK1(sK0),u1_struct_0(sK0)) [avatar definition]
289. ~m1_subset_1(sK1(sK0),u1_struct_0(sK0)) <- (~16) [avatar component clause 287] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
291. 17 <=> m1_subset_1(sK2(sK0),u1_struct_0(sK0)) [avatar definition]
293. ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) <- (~17) [avatar component clause 291] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
295. 18 <=> m1_subset_1(sK3(sK0),u1_struct_0(sK0)) [avatar definition]
297. ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) <- (~18) [avatar component clause 295] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
299. 19 <=> v2_struct_0(sK0) [avatar definition]
301. v2_struct_0(sK0) <- (19) [avatar component clause 299] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
302. ~16 | ~17 | ~18 | 19 | ~1 | ~2 | 3 [avatar split clause 285,164,159,154,299,295,291,287]
304. ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~16) [resolution 289,135] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
306. ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (1, ~16) [subsumption resolution 304,156] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
307. v7_struct_0(sK0) <- (1, 2, ~16) [subsumption resolution 306,161] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
308. $false <- (1, 2, ~3, ~16) [subsumption resolution 307,166] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
309. ~1 | ~2 | 3 | 16 [avatar contradiction clause 308]
321. ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~17) [resolution 293,136] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
323. ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (1, ~17) [subsumption resolution 321,156] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
324. v7_struct_0(sK0) <- (1, 2, ~17) [subsumption resolution 323,161] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
325. $false <- (1, 2, ~3, ~17) [subsumption resolution 324,166] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
326. ~1 | ~2 | 3 | 17 [avatar contradiction clause 325]
334. ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~18) [resolution 297,137] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
336. ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (1, ~18) [subsumption resolution 334,156] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
337. v7_struct_0(sK0) <- (1, 2, ~18) [subsumption resolution 336,161] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
338. $false <- (1, 2, ~3, ~18) [subsumption resolution 337,166] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
339. ~1 | ~2 | 3 | 18 [avatar contradiction clause 338]
341. v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (19) [resolution 301,119] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
344. ~l1_struct_0(sK0) <- (~3, 19) [subsumption resolution 341,166] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
345. $false <- (~3, 10, 19) [subsumption resolution 344,203] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
346. 3 | ~10 | ~19 [avatar contradiction clause 345]
348. $false [avatar sat refutation 157,162,167,204,302,309,326,339,346] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:189,thDist:-189}
% SZS output end Proof for diraf__t36_diraf
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 44
% Generated clauses: 149
% Active clauses: 56
% Passive clauses: 67
% Final active clauses: 53
% Final passive clauses: 6
% Main loop iterations started: 61
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 18
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 43
% Unit resulting resolution: 4
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 12
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 9
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% sat proof minimization: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

