i: [1,139,0,2,1,1,0,0] t18_afvect0
a: [3,223,0,2,1,139]
s: 223
i: [1,138,0,2,1,1,0,0] t18_afvect0
a: [3,228,0,2,1,138]
s: 228
i: [1,137,0,2,1,1,0,0] t18_afvect0
a: [3,233,0,2,1,137]
s: 233
i: [1,144,0,9,2,1,0,0] t18_afvect0
a: [3,201,0,4,1,144]
s: 201
i: [1,141,0,4,1,1,0,0] t18_afvect0
a: [3,213,0,4,1,141]
s: 213
i: [1,140,0,4,1,1,0,0] t18_afvect0
a: [3,218,0,4,1,140]
s: 218
a: [3,198,0,5,1,144]
s: 198
i: [1,145,0,2,1,0,0,255] spc1_boole
a: [3,238,0,2,1,145]
s: 238
i: [1,146,0,2,1,0,0,2] fc1_xboole_0
a: [3,243,0,2,1,146]
s: 243
i: [1,186,0,2,1,0,0,2] rc2_xboole_0
a: [3,248,0,2,1,186]
s: 248
i: [1,187,0,2,1,0,0,1] existence_l1_analoaf
a: [3,253,0,2,1,187]
s: 253
i: [1,188,0,2,1,0,0,2] existence_l1_struct_0
a: [3,258,0,2,1,188]
s: 258
i: [1,189,0,2,1,0,0,2] rc1_xboole_0
a: [3,263,0,2,1,189]
s: 263
i: [1,191,0,2,1,0,0,255] rc1_relat_1
a: [3,268,0,2,1,191]
s: 268
i: [1,190,0,2,1,0,0,255] rc1_relat_1
a: [3,273,0,2,1,190]
s: 273
i: [1,192,0,2,1,0,0,255] rc2_relat_1
a: [3,278,0,2,1,192]
s: 278
i: [1,147,0,4,2,0,0,1] dt_l1_analoaf
s: 147
i: [1,162,0,4,2,0,0,255] cc1_relat_1
s: 162
d: [2,280,1,2,1,1,40,147,223]
a: [3,285,1,2,1,280]
s: 285
i: [1,182,0,4,1,0,0,1] existence_m1_subset_1
s: 182
i: [1,163,0,5,2,0,0,1] t6_boole
s: 163
d: [2,281,1,2,1,1,40,147,253]
a: [3,290,1,2,1,281]
s: 290
i: [1,181,0,5,2,0,0,255] t7_boole
s: 181
d: [2,292,1,2,1,1,40,162,243]
a: [3,297,1,2,1,292]
s: 297
i: [1,148,0,6,3,0,0,255] cc7_struct_0
s: 148
i: [1,149,0,6,3,0,0,255] cc5_struct_0
s: 149
d: [2,293,1,2,1,1,40,162,263]
a: [3,302,1,2,1,293]
s: 302
i: [1,151,0,6,3,0,0,3] cc17_struct_0
s: 151
d: [2,307,1,4,2,1,40,148,258]
a: [3,313,1,2,1,307]
s: 313
i: [1,152,0,6,3,0,0,3] cc2_struct_0
s: 152
i: [1,153,0,6,3,0,0,3] cc1_struct_0
s: 153
d: [2,309,2,4,2,1,40,148,290]
a: [3,322,2,2,1,309]
s: 322
i: [1,156,0,6,3,0,0,255] cc4_struct_0
s: 156
i: [1,157,0,6,3,0,0,255] cc6_struct_0
s: 157
i: [1,178,0,6,2,0,0,255] t1_subset
s: 178
i: [1,183,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 183
i: [1,154,0,7,3,0,0,2] cc8_struct_0
s: 154
i: [1,158,0,7,3,0,0,255] cc11_struct_0
s: 158
i: [1,159,0,7,3,0,0,255] cc11_struct_0
s: 159
i: [1,160,0,7,3,0,0,2] cc9_struct_0
s: 160
i: [1,164,0,7,3,0,0,3] fc2_struct_0
s: 164
i: [1,165,0,7,3,0,0,255] fc9_struct_0
s: 165
i: [1,172,0,7,3,0,0,255] fc6_struct_0
s: 172
i: [1,173,0,7,3,0,0,3] fc1_struct_0
s: 173
i: [1,176,0,7,3,0,0,255] fc8_struct_0
s: 176
i: [1,177,0,7,3,0,0,255] fc7_struct_0
s: 177
i: [1,180,0,7,3,0,0,1] t8_boole
s: 180
i: [1,179,0,8,3,0,0,255] t2_subset
s: 179
d: [2,339,1,7,2,1,40,179,213]
a: [3,349,1,3,1,339]
s: 349
r: 349
i: [1,161,0,9,4,0,0,255] cc10_struct_0
s: 161
i: [1,193,0,18,6,0,0,2] d2_afvect0
d: [2,194,0,14,5,0,18,193]
s: 194
a: [3,348,1,3,1,339]
s: 348
i: [1,184,0,18,6,0,0,1] irreflexivity_r1_afvect0
s: 184
d: [2,356,2,4,2,1,40,349,164]
d: [2,358,2,2,1,2,21,356,285]
a: [3,361,2,2,1,358]
s: 361
i: [1,166,0,20,6,0,0,2] t6_afvect0
s: 166
i: [1,185,0,22,7,0,0,1] symmetry_r1_afvect0
s: 185
r: 198
r: 201
i: [1,143,0,8,2,1,0,0] t18_afvect0
a: [3,207,0,3,1,143]
s: 207
a: [3,202,0,4,1,144]
s: 202
a: [3,345,1,4,1,339]
s: 345
a: [3,197,0,5,1,144]
s: 197
i: [1,167,0,24,7,0,0,2] d2_afvect0
s: 167
d: [2,340,1,7,2,1,40,179,218]
a: [3,354,1,4,1,340]
s: 354
i: [1,169,0,27,8,0,0,2] d2_afvect0
s: 169
d: [2,387,2,4,1,1,40,345,183]
a: [3,393,2,4,1,387]
s: 393
i: [1,170,0,29,8,0,0,1] d3_afvect0
s: 170
i: [1,171,0,29,8,0,0,1] d3_afvect0
s: 171
d: [2,395,2,4,1,1,40,354,183]
a: [3,401,2,4,1,395]
s: 401
d: [2,341,1,6,2,0,40,179,182]
s: 341
d: [2,415,2,6,2,0,40,341,183]
s: 415
d: [2,304,1,7,3,0,44,163,163]
s: 304
d: [2,333,1,7,3,0,40,173,162]
s: 333
d: [2,328,1,9,4,0,44,154,163]
s: 328
d: [2,330,1,9,4,0,44,160,163]
s: 330
d: [2,409,1,37,12,0,40,171,166]
d: [2,414,1,15,5,0,18,409]
s: 414
d: [2,410,1,41,13,0,40,171,167]
d: [2,413,1,23,7,0,18,410]
s: 413
d: [2,407,1,44,14,0,40,170,169]
d: [2,408,1,26,8,0,18,407]
s: 408
d: [2,430,2,0,0,8,52,223,228,233,207,202,213,218,197,408]
e: 430
d: [2,431,2,21,7,1,40,408,197]
d: [2,437,2,17,6,2,21,431,218]
d: [2,438,2,13,5,3,21,437,213]
d: [2,439,2,11,4,4,21,438,223]
d: [2,440,2,9,3,5,21,439,228]
d: [2,441,2,7,2,6,21,440,233]
d: [2,442,2,4,1,7,21,441,207]
d: [2,443,2,0,0,8,21,442,202]
e: 443
r: 197
r: 207
a: [3,206,0,3,1,143]
s: 206
r: 393
r: 345
r: 202
r: 213
d: [2,447,0,4,1,2,23,202,206]
a: [3,452,0,4,1,447]
s: 452
s: 198
d: [2,454,0,5,1,2,22,198,206]
a: [3,458,0,5,1,454]
s: 458
d: [2,460,2,0,0,5,52,233,228,223,218,458,414]
e: 460
d: [2,461,2,10,4,1,40,458,414]
d: [2,465,2,6,3,2,21,461,218]
d: [2,466,2,4,2,3,21,465,223]
d: [2,467,2,2,1,4,21,466,228]
d: [2,468,2,0,0,5,21,467,233]
e: 468
r: 206
r: 458
s: 207
s: 201
s: 213
s: 345
s: 198
d: [2,483,2,0,0,7,52,223,233,228,201,213,218,198,413]
e: 483
d: [2,484,2,18,6,1,40,198,413]
d: [2,486,2,14,5,2,21,484,218]
d: [2,487,2,10,4,3,21,486,213]
d: [2,488,2,8,3,4,21,487,223]
d: [2,489,2,6,2,5,21,488,228]
d: [2,490,2,4,1,6,21,489,233]
d: [2,491,2,0,0,7,21,490,201]
e: 491
i: [1,142,0,12,3,1,0,0] t18_afvect0
f: 491,468,454,443,137,138,139,140,141,142,143,144
% Refutation found. Thanks to Tanya!
% SZS status Theorem for afvect0__t18_afvect0
% SZS output start Proof for afvect0__t18_afvect0
1. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_afvect0(X0,X1,X2,X1) <=> (r1_afvect0(X0,X1,X2) | X1 = X2))))) [input t18_afvect0]
2. ~! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_afvect0(X0,X1,X2,X1) <=> (r1_afvect0(X0,X1,X2) | X1 = X2))))) [negated conjecture 1]
18. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r1_afvect0(X0,X1,X2) <=> (X1 != X2 & r2_analoaf(X0,X1,X2,X2,X1)))))) [input d2_afvect0]
19. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r2_afvect0(X0,X1,X2,X3) <=> r2_analoaf(X0,X1,X2,X2,X3)))))) [input d3_afvect0]
44. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => r2_analoaf(X0,X1,X1,X2,X2)))) [input t6_afvect0]
55. ? [X0] : (? [X1] : (? [X2] : ((r2_afvect0(X0,X1,X2,X1) <~> (r1_afvect0(X0,X1,X2) | X1 = X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
56. ? [X0] : (? [X1] : (? [X2] : ((r2_afvect0(X0,X1,X2,X1) <~> (r1_afvect0(X0,X1,X2) | X1 = X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) [flattening 55]
86. ! [X0] : (! [X1] : (! [X2] : (r2_analoaf(X0,X1,X1,X2,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 44]
87. ! [X0] : (! [X1] : (! [X2] : (r2_analoaf(X0,X1,X1,X2,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 86]
88. ! [X0] : (! [X1] : (! [X2] : ((r1_afvect0(X0,X1,X2) <=> (X1 != X2 & r2_analoaf(X0,X1,X2,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 18]
89. ! [X0] : (! [X1] : (! [X2] : ((r1_afvect0(X0,X1,X2) <=> (X1 != X2 & r2_analoaf(X0,X1,X2,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 88]
90. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_afvect0(X0,X1,X2,X3) <=> r2_analoaf(X0,X1,X2,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 19]
91. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_afvect0(X0,X1,X2,X3) <=> r2_analoaf(X0,X1,X2,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 90]
114. ? [X0] : (? [X1] : (? [X2] : ((((~r1_afvect0(X0,X1,X2) & X1 != X2) | ~r2_afvect0(X0,X1,X2,X1)) & ((r1_afvect0(X0,X1,X2) | X1 = X2) | r2_afvect0(X0,X1,X2,X1))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) [nnf transformation 56]
115. ? [X0] : (? [X1] : (? [X2] : (((~r1_afvect0(X0,X1,X2) & X1 != X2) | ~r2_afvect0(X0,X1,X2,X1)) & (r1_afvect0(X0,X1,X2) | X1 = X2 | r2_afvect0(X0,X1,X2,X1)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) [flattening 114]
116. ? [X0] : (? [X1] : (? [X2] : (((~r1_afvect0(X0,X1,X2) & X1 != X2) | ~r2_afvect0(X0,X1,X2,X1)) & (r1_afvect0(X0,X1,X2) | X1 = X2 | r2_afvect0(X0,X1,X2,X1)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : (((~r1_afvect0(sK0,X1,X2) & X1 != X2) | ~r2_afvect0(sK0,X1,X2,X1)) & (r1_afvect0(sK0,X1,X2) | X1 = X2 | r2_afvect0(sK0,X1,X2,X1)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0)) [choice axiom]
117. ? [X1] : (? [X2] : (((~r1_afvect0(sK0,X1,X2) & X1 != X2) | ~r2_afvect0(sK0,X1,X2,X1)) & (r1_afvect0(sK0,X1,X2) | X1 = X2 | r2_afvect0(sK0,X1,X2,X1)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (((~r1_afvect0(sK0,sK1,X2) & sK1 != X2) | ~r2_afvect0(sK0,sK1,X2,sK1)) & (r1_afvect0(sK0,sK1,X2) | sK1 = X2 | r2_afvect0(sK0,sK1,X2,sK1)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
118. ? [X2] : (((~r1_afvect0(sK0,sK1,X2) & sK1 != X2) | ~r2_afvect0(sK0,sK1,X2,sK1)) & (r1_afvect0(sK0,sK1,X2) | sK1 = X2 | r2_afvect0(sK0,sK1,X2,sK1)) & m1_subset_1(X2,u1_struct_0(sK0))) => (((~r1_afvect0(sK0,sK1,sK2) & sK1 != sK2) | ~r2_afvect0(sK0,sK1,sK2,sK1)) & (r1_afvect0(sK0,sK1,sK2) | sK1 = sK2 | r2_afvect0(sK0,sK1,sK2,sK1)) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
119. ((((~r1_afvect0(sK0,sK1,sK2) & sK1 != sK2) | ~r2_afvect0(sK0,sK1,sK2,sK1)) & (r1_afvect0(sK0,sK1,sK2) | sK1 = sK2 | r2_afvect0(sK0,sK1,sK2,sK1)) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0) [skolemisation 115,118,117,116]
120. ! [X0] : (! [X1] : (! [X2] : (((r1_afvect0(X0,X1,X2) | (X1 = X2 | ~r2_analoaf(X0,X1,X2,X2,X1))) & ((X1 != X2 & r2_analoaf(X0,X1,X2,X2,X1)) | ~r1_afvect0(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [nnf transformation 89]
121. ! [X0] : (! [X1] : (! [X2] : (((r1_afvect0(X0,X1,X2) | X1 = X2 | ~r2_analoaf(X0,X1,X2,X2,X1)) & ((X1 != X2 & r2_analoaf(X0,X1,X2,X2,X1)) | ~r1_afvect0(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 120]
122. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r2_afvect0(X0,X1,X2,X3) | ~r2_analoaf(X0,X1,X2,X2,X3)) & (r2_analoaf(X0,X1,X2,X2,X3) | ~r2_afvect0(X0,X1,X2,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [nnf transformation 91]
137. ~v7_struct_0(sK0) [cnf transformation 119] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
138. v1_afvect0(sK0) [cnf transformation 119] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
139. l1_analoaf(sK0) [cnf transformation 119] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
140. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 119] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
141. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 119] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
142. r1_afvect0(sK0,sK1,sK2) | sK1 = sK2 | r2_afvect0(sK0,sK1,sK2,sK1) [cnf transformation 119] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
143. sK1 != sK2 | ~r2_afvect0(sK0,sK1,sK2,sK1) [cnf transformation 119] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
144. ~r1_afvect0(sK0,sK1,sK2) | ~r2_afvect0(sK0,sK1,sK2,sK1) [cnf transformation 119] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
166. r2_analoaf(X0,X1,X1,X2,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 87] {a:0,w:20,wCS:8000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
167. r2_analoaf(X0,X1,X2,X2,X1) | ~r1_afvect0(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 121] {a:0,w:24,wCS:9600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
169. ~r2_analoaf(X0,X1,X2,X2,X1) | X1 = X2 | r1_afvect0(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 121] {a:0,w:27,wCS:10800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
170. r2_analoaf(X0,X1,X2,X2,X3) | ~r2_afvect0(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 122] {a:0,w:29,wCS:11600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
171. ~r2_analoaf(X0,X1,X2,X2,X3) | r2_afvect0(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 122] {a:0,w:29,wCS:11600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
196. 1 <=> r2_afvect0(sK0,sK1,sK2,sK1) [avatar definition]
197. r2_afvect0(sK0,sK1,sK2,sK1) <- (1) [avatar component clause 196] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
198. ~r2_afvect0(sK0,sK1,sK2,sK1) <- (~1) [avatar component clause 196] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
200. 2 <=> r1_afvect0(sK0,sK1,sK2) [avatar definition]
201. r1_afvect0(sK0,sK1,sK2) <- (2) [avatar component clause 200] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
202. ~r1_afvect0(sK0,sK1,sK2) <- (~2) [avatar component clause 200] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
203. ~1 | ~2 [avatar split clause 144,200,196]
205. 3 <=> sK1 = sK2 [avatar definition]
206. sK1 = sK2 <- (3) [avatar component clause 205] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
207. sK1 != sK2 <- (~3) [avatar component clause 205] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
208. ~1 | ~3 [avatar split clause 143,205,196]
209. 1 | 3 | 2 [avatar split clause 142,200,205,196]
211. 4 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
213. m1_subset_1(sK2,u1_struct_0(sK0)) <- (4) [avatar component clause 211] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
214. 4 [avatar split clause 141,211]
216. 5 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
218. m1_subset_1(sK1,u1_struct_0(sK0)) <- (5) [avatar component clause 216] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
219. 5 [avatar split clause 140,216]
221. 6 <=> l1_analoaf(sK0) [avatar definition]
223. l1_analoaf(sK0) <- (6) [avatar component clause 221] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
224. 6 [avatar split clause 139,221]
226. 7 <=> v1_afvect0(sK0) [avatar definition]
228. v1_afvect0(sK0) <- (7) [avatar component clause 226] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
229. 7 [avatar split clause 138,226]
231. 8 <=> v7_struct_0(sK0) [avatar definition]
233. ~v7_struct_0(sK0) <- (~8) [avatar component clause 231] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
234. ~8 [avatar split clause 137,231]
407. ~r2_afvect0(X0,X1,X2,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) | X1 = X2 | r1_afvect0(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [resolution 170,169] {a:1,w:44,wCS:17600,sine:1,thAx:0,allAx:2,thDist:-2}
408. ~r2_afvect0(X0,X1,X2,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) | X1 = X2 | r1_afvect0(X0,X1,X2) [duplicate literal removal 407] {a:1,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
409. r2_afvect0(X0,X1,X1,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [resolution 171,166] {a:1,w:37,wCS:14800,sine:1,thAx:0,allAx:2,thDist:-2}
410. r2_afvect0(X2,X3,X4,X3) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~m1_subset_1(X4,u1_struct_0(X2)) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l1_analoaf(X2) | ~v1_afvect0(X2) | v7_struct_0(X2) | ~r1_afvect0(X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X2)) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l1_analoaf(X2) | ~v1_afvect0(X2) | v7_struct_0(X2) [resolution 171,167] {a:1,w:41,wCS:16400,sine:1,thAx:0,allAx:2,thDist:-2}
413. r2_afvect0(X2,X3,X4,X3) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~m1_subset_1(X4,u1_struct_0(X2)) | ~l1_analoaf(X2) | ~v1_afvect0(X2) | v7_struct_0(X2) | ~r1_afvect0(X2,X3,X4) [duplicate literal removal 410] {a:1,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
414. r2_afvect0(X0,X1,X1,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [duplicate literal removal 409] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
431. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) | sK1 = sK2 | r1_afvect0(sK0,sK1,sK2) <- (1) [resolution 408,197] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
437. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) | sK1 = sK2 | r1_afvect0(sK0,sK1,sK2) <- (1, 5) [subsumption resolution 431,218] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
438. ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) | sK1 = sK2 | r1_afvect0(sK0,sK1,sK2) <- (1, 4, 5) [subsumption resolution 437,213] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
439. ~v1_afvect0(sK0) | v7_struct_0(sK0) | sK1 = sK2 | r1_afvect0(sK0,sK1,sK2) <- (1, 4, 5, 6) [subsumption resolution 438,223] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
440. v7_struct_0(sK0) | sK1 = sK2 | r1_afvect0(sK0,sK1,sK2) <- (1, 4, 5, 6, 7) [subsumption resolution 439,228] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
441. sK1 = sK2 | r1_afvect0(sK0,sK1,sK2) <- (1, 4, 5, 6, 7, ~8) [subsumption resolution 440,233] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
442. r1_afvect0(sK0,sK1,sK2) <- (1, ~3, 4, 5, 6, 7, ~8) [subsumption resolution 441,207] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
443. $false <- (1, ~2, ~3, 4, 5, 6, 7, ~8) [subsumption resolution 442,202] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
444. ~1 | 2 | 3 | ~4 | ~5 | ~6 | ~7 | 8 [avatar contradiction clause 443]
454. ~r2_afvect0(sK0,sK1,sK1,sK1) <- (~1, 3) [forward demodulation 198,206] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
456. 35 <=> r2_afvect0(sK0,sK1,sK1,sK1) [avatar definition]
458. ~r2_afvect0(sK0,sK1,sK1,sK1) <- (~35) [avatar component clause 456] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
459. ~35 | 1 | ~3 [avatar split clause 454,205,196,456]
461. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~35) [resolution 458,414] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
465. ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (5, ~35) [subsumption resolution 461,218] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
466. ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (5, 6, ~35) [subsumption resolution 465,223] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
467. v7_struct_0(sK0) <- (5, 6, 7, ~35) [subsumption resolution 466,228] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
468. $false <- (5, 6, 7, ~8, ~35) [subsumption resolution 467,233] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
469. ~5 | ~6 | ~7 | 8 | 35 [avatar contradiction clause 468]
484. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) | ~r1_afvect0(sK0,sK1,sK2) <- (~1) [resolution 198,413] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
486. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) | ~r1_afvect0(sK0,sK1,sK2) <- (~1, 5) [subsumption resolution 484,218] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
487. ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) | ~r1_afvect0(sK0,sK1,sK2) <- (~1, 4, 5) [subsumption resolution 486,213] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
488. ~v1_afvect0(sK0) | v7_struct_0(sK0) | ~r1_afvect0(sK0,sK1,sK2) <- (~1, 4, 5, 6) [subsumption resolution 487,223] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
489. v7_struct_0(sK0) | ~r1_afvect0(sK0,sK1,sK2) <- (~1, 4, 5, 6, 7) [subsumption resolution 488,228] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
490. ~r1_afvect0(sK0,sK1,sK2) <- (~1, 4, 5, 6, 7, ~8) [subsumption resolution 489,233] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
491. $false <- (~1, 2, 4, 5, 6, 7, ~8) [subsumption resolution 490,201] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
492. 1 | ~2 | ~4 | ~5 | ~6 | ~7 | 8 [avatar contradiction clause 491]
493. $false [avatar sat refutation 203,208,209,214,219,224,229,234,444,459,469,492] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:319,thDist:-319}
% SZS output end Proof for afvect0__t18_afvect0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 46
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 56
% Generated clauses: 217
% Active clauses: 78
% Passive clauses: 88
% Final active clauses: 67
% Final passive clauses: 3
% Main loop iterations started: 85
% 
% >>> Simplifying Inferences
% Duplicate literals: 63
% Fw subsumption resolutions: 29
% Fw demodulations: 1
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Equational tautologies: 2
% Forward subsumptions: 22
% 
% >>> Generating Inferences
% Binary resolution: 53
% Unit resulting resolution: 3
% Forward superposition: 3
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 17
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 44
% SAT solver unit clauses: 14
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% minimizing solver time: 0.001 s
% congruence closure: 0.001 s ( own 0.0000 s ) 
% other: 0.003 s ( own 0.002 s ) 

