i: [1,78,0,2,1,1,0,0] t33_e_siec
a: [3,124,0,2,1,78]
s: 124
i: [1,77,0,2,1,1,0,0] t33_e_siec
a: [3,129,0,2,1,77]
s: 129
i: [1,76,0,2,1,1,0,0] t33_e_siec
a: [3,134,0,2,1,76]
s: 134
i: [1,107,0,2,1,0,0,1] existence_l1_e_siec
a: [3,139,0,2,1,107]
s: 139
i: [1,109,0,2,1,0,0,1] rc2_e_siec
a: [3,144,0,2,1,109]
s: 144
i: [1,108,0,2,1,0,0,1] rc2_e_siec
a: [3,149,0,2,1,108]
s: 149
i: [1,111,0,2,1,0,0,1] rc3_e_siec
a: [3,154,0,2,1,111]
s: 154
i: [1,110,0,2,1,0,0,1] rc3_e_siec
a: [3,159,0,2,1,110]
s: 159
i: [1,80,0,3,1,0,0,1] dt_k4_relat_1
s: 80
i: [1,83,0,3,1,0,0,255] fc28_relat_1
s: 83
i: [1,99,0,3,1,0,0,1] reflexivity_r1_tarski
s: 99
i: [1,79,0,26,2,1,0,0] t33_e_siec
a: [3,115,0,13,1,79]
s: 115
i: [1,84,0,4,1,0,0,255] fc28_relat_1
s: 84
i: [1,85,0,4,1,0,0,255] fc28_relat_1
s: 85
i: [1,98,0,4,1,0,0,255] existence_m1_subset_1
s: 98
i: [1,100,0,4,1,0,0,1] fc6_relat_1
s: 100
i: [1,104,0,4,1,0,0,255] rc3_relat_1
s: 104
i: [1,81,0,5,1,0,0,255] rd2_relat_1
s: 81
i: [1,86,0,5,2,0,0,1] dt_u2_e_siec
s: 86
i: [1,87,0,5,2,0,0,1] dt_u1_e_siec
s: 87
i: [1,88,0,5,2,0,0,1] dt_k2_relat_1
s: 88
i: [1,105,0,5,1,0,0,255] rc3_relat_1
s: 105
i: [1,106,0,5,1,0,0,255] rc3_relat_1
s: 106
i: [1,82,0,6,1,0,0,1] rd3_relat_1
s: 82
i: [1,101,0,6,2,0,0,1] fc2_relat_1
s: 101
i: [1,89,0,7,2,0,0,1] involutiveness_k2_relat_1
s: 89
i: [1,95,0,7,2,0,0,255] t3_subset
s: 95
i: [1,96,0,7,2,0,0,255] t3_subset
s: 96
d: [2,164,1,5,1,0,40,95,98]
s: 164
i: [1,90,0,8,3,0,0,255] cc2_relat_1
s: 90
i: [1,91,0,12,4,0,0,255] d9_e_siec
s: 91
d: [2,166,1,6,2,0,40,90,98]
s: 166
i: [1,102,0,12,4,0,0,255] cc5_relat_1
s: 102
d: [2,167,1,7,3,0,40,90,96]
s: 167
i: [1,103,0,12,4,0,0,255] cc6_relat_1
s: 103
i: [1,97,0,13,3,0,0,1] t4_sysrel
s: 97
d: [2,174,1,19,2,1,40,97,115]
a: [3,181,1,7,1,174]
s: 181
i: [1,92,0,14,4,0,0,1] d10_e_siec
s: 92
d: [2,189,2,3,1,1,40,181,101]
a: [3,193,2,3,1,189]
s: 193
d: [2,196,3,0,0,2,52,124,193,87]
e: 196
d: [2,197,3,2,1,1,40,193,87]
d: [2,199,3,0,0,2,21,197,124]
e: 199
r: 181
r: 193
i: [1,93,0,18,4,0,0,1] t18_e_siec
s: 93
i: [1,94,0,18,4,0,0,1] t18_e_siec
s: 94
a: [3,192,2,3,1,189]
s: 192
a: [3,180,1,7,1,174]
s: 180
a: [3,185,1,12,1,174]
s: 185
d: [2,205,2,0,0,4,52,134,129,124,185,93]
e: 205
d: [2,206,2,6,3,1,40,185,93]
d: [2,208,2,4,2,2,21,206,124]
d: [2,209,2,2,1,3,21,208,129]
d: [2,210,2,0,0,4,21,209,134]
e: 210
r: 115
r: 185
a: [3,114,0,13,1,79]
s: 114
a: [3,119,0,13,1,79]
s: 119
d: [2,219,1,19,2,1,40,119,97]
a: [3,223,1,7,1,219]
s: 223
a: [3,184,1,12,1,174]
s: 184
d: [2,229,2,3,1,1,40,223,101]
a: [3,233,2,3,1,229]
s: 233
d: [2,236,3,0,0,2,52,124,233,86]
e: 236
d: [2,237,3,2,1,1,40,233,86]
d: [2,239,3,0,0,2,21,237,124]
e: 239
r: 223
r: 233
a: [3,222,1,7,1,219]
s: 222
a: [3,227,1,12,1,219]
s: 227
d: [2,241,2,0,0,4,52,134,129,124,227,94]
e: 241
d: [2,242,2,6,3,1,40,227,94]
d: [2,244,2,4,2,2,21,242,124]
d: [2,245,2,2,1,3,21,244,129]
d: [2,246,2,0,0,4,21,245,134]
e: 246
f: 246,239,229,219,210,199,189,174,76,77,78,79
% Refutation found. Thanks to Tanya!
% SZS status Theorem for e_siec__t33_e_siec
% SZS output start Proof for e_siec__t33_e_siec
1. ! [X0] : ((l1_e_siec(X0) & v3_e_siec(X0) & v2_e_siec(X0)) => (r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0))) & r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0))))) [input t33_e_siec]
2. ~! [X0] : ((l1_e_siec(X0) & v3_e_siec(X0) & v2_e_siec(X0)) => (r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0))) & r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0))))) [negated conjecture 1]
19. ! [X0] : (l1_e_siec(X0) => v1_relat_1(u1_e_siec(X0))) [input dt_u1_e_siec]
21. ! [X0] : (l1_e_siec(X0) => v1_relat_1(u2_e_siec(X0))) [input dt_u2_e_siec]
26. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k4_xboole_0(X0,X1))) [input fc2_relat_1]
35. ! [X0] : ((l1_e_siec(X0) & v3_e_siec(X0) & v2_e_siec(X0)) => (r1_tarski(k4_xboole_0(u2_e_siec(X0),k4_relat_1(u1_struct_0(X0))),k2_zfmisc_1(k8_e_siec(X0),k7_e_siec(X0))) & r1_tarski(k4_xboole_0(u1_e_siec(X0),k4_relat_1(u1_struct_0(X0))),k2_zfmisc_1(k8_e_siec(X0),k7_e_siec(X0))))) [input t18_e_siec]
37. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X2,k2_zfmisc_1(X0,X1)) => r1_tarski(k2_relat_1(X2),k2_zfmisc_1(X1,X0)))) [input t4_sysrel]
39. ! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(X2,k2_zfmisc_1(X0,X1)) => r1_tarski(k2_relat_1(X2),k2_zfmisc_1(X1,X0)))) [flattening 37]
43. ? [X0] : ((~r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0))) | ~r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0)))) & (l1_e_siec(X0) & v3_e_siec(X0) & v2_e_siec(X0))) [ennf transformation 2]
44. ? [X0] : ((~r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0))) | ~r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0)))) & l1_e_siec(X0) & v3_e_siec(X0) & v2_e_siec(X0)) [flattening 43]
45. ! [X0] : (v1_relat_1(u2_e_siec(X0)) | ~l1_e_siec(X0)) [ennf transformation 21]
46. ! [X0] : (v1_relat_1(u1_e_siec(X0)) | ~l1_e_siec(X0)) [ennf transformation 19]
54. ! [X0] : ((r1_tarski(k4_xboole_0(u2_e_siec(X0),k4_relat_1(u1_struct_0(X0))),k2_zfmisc_1(k8_e_siec(X0),k7_e_siec(X0))) & r1_tarski(k4_xboole_0(u1_e_siec(X0),k4_relat_1(u1_struct_0(X0))),k2_zfmisc_1(k8_e_siec(X0),k7_e_siec(X0)))) | (~l1_e_siec(X0) | ~v3_e_siec(X0) | ~v2_e_siec(X0))) [ennf transformation 35]
55. ! [X0] : ((r1_tarski(k4_xboole_0(u2_e_siec(X0),k4_relat_1(u1_struct_0(X0))),k2_zfmisc_1(k8_e_siec(X0),k7_e_siec(X0))) & r1_tarski(k4_xboole_0(u1_e_siec(X0),k4_relat_1(u1_struct_0(X0))),k2_zfmisc_1(k8_e_siec(X0),k7_e_siec(X0)))) | ~l1_e_siec(X0) | ~v3_e_siec(X0) | ~v2_e_siec(X0)) [flattening 54]
56. ! [X0,X1,X2] : ((r1_tarski(k2_relat_1(X2),k2_zfmisc_1(X1,X0)) | ~r1_tarski(X2,k2_zfmisc_1(X0,X1))) | ~v1_relat_1(X2)) [ennf transformation 39]
57. ! [X0,X1,X2] : (r1_tarski(k2_relat_1(X2),k2_zfmisc_1(X1,X0)) | ~r1_tarski(X2,k2_zfmisc_1(X0,X1)) | ~v1_relat_1(X2)) [flattening 56]
58. ! [X0,X1] : (v1_relat_1(k4_xboole_0(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 26]
63. ? [X0] : ((~r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0))) | ~r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(X0),k4_relat_1(u1_struct_0(X0)))),k2_zfmisc_1(k7_e_siec(X0),k8_e_siec(X0)))) & l1_e_siec(X0) & v3_e_siec(X0) & v2_e_siec(X0)) => ((~r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0))) | ~r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0)))) & l1_e_siec(sK0) & v3_e_siec(sK0) & v2_e_siec(sK0)) [choice axiom]
64. (~r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0))) | ~r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0)))) & l1_e_siec(sK0) & v3_e_siec(sK0) & v2_e_siec(sK0) [skolemisation 44,63]
76. v2_e_siec(sK0) [cnf transformation 64] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
77. v3_e_siec(sK0) [cnf transformation 64] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
78. l1_e_siec(sK0) [cnf transformation 64] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
79. ~r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0))) | ~r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0))) [cnf transformation 64] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
86. v1_relat_1(u2_e_siec(X0)) | ~l1_e_siec(X0) [cnf transformation 45] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. v1_relat_1(u1_e_siec(X0)) | ~l1_e_siec(X0) [cnf transformation 46] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
93. r1_tarski(k4_xboole_0(u1_e_siec(X0),k4_relat_1(u1_struct_0(X0))),k2_zfmisc_1(k8_e_siec(X0),k7_e_siec(X0))) | ~l1_e_siec(X0) | ~v3_e_siec(X0) | ~v2_e_siec(X0) [cnf transformation 55] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
94. r1_tarski(k4_xboole_0(u2_e_siec(X0),k4_relat_1(u1_struct_0(X0))),k2_zfmisc_1(k8_e_siec(X0),k7_e_siec(X0))) | ~l1_e_siec(X0) | ~v3_e_siec(X0) | ~v2_e_siec(X0) [cnf transformation 55] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
97. r1_tarski(k2_relat_1(X2),k2_zfmisc_1(X1,X0)) | ~r1_tarski(X2,k2_zfmisc_1(X0,X1)) | ~v1_relat_1(X2) [cnf transformation 57] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. v1_relat_1(k4_xboole_0(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 58] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
113. 1 <=> r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0))) [avatar definition]
115. ~r1_tarski(k2_relat_1(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0))) <- (~1) [avatar component clause 113] {a:0,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
117. 2 <=> r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0))) [avatar definition]
119. ~r1_tarski(k2_relat_1(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))),k2_zfmisc_1(k7_e_siec(sK0),k8_e_siec(sK0))) <- (~2) [avatar component clause 117] {a:0,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
120. ~1 | ~2 [avatar split clause 79,117,113]
122. 3 <=> l1_e_siec(sK0) [avatar definition]
124. l1_e_siec(sK0) <- (3) [avatar component clause 122] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
125. 3 [avatar split clause 78,122]
127. 4 <=> v3_e_siec(sK0) [avatar definition]
129. v3_e_siec(sK0) <- (4) [avatar component clause 127] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
130. 4 [avatar split clause 77,127]
132. 5 <=> v2_e_siec(sK0) [avatar definition]
134. v2_e_siec(sK0) <- (5) [avatar component clause 132] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
135. 5 [avatar split clause 76,132]
174. ~r1_tarski(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0))),k2_zfmisc_1(k8_e_siec(sK0),k7_e_siec(sK0))) | ~v1_relat_1(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))) <- (~1) [resolution 97,115] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
179. 11 <=> v1_relat_1(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))) [avatar definition]
181. ~v1_relat_1(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))) <- (~11) [avatar component clause 179] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
183. 12 <=> r1_tarski(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0))),k2_zfmisc_1(k8_e_siec(sK0),k7_e_siec(sK0))) [avatar definition]
185. ~r1_tarski(k4_xboole_0(u1_e_siec(sK0),k4_relat_1(u1_struct_0(sK0))),k2_zfmisc_1(k8_e_siec(sK0),k7_e_siec(sK0))) <- (~12) [avatar component clause 183] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
186. ~11 | ~12 | 1 [avatar split clause 174,113,183,179]
189. ~v1_relat_1(u1_e_siec(sK0)) <- (~11) [resolution 181,101] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
191. 13 <=> v1_relat_1(u1_e_siec(sK0)) [avatar definition]
193. ~v1_relat_1(u1_e_siec(sK0)) <- (~13) [avatar component clause 191] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
194. ~13 | 11 [avatar split clause 189,179,191]
197. ~l1_e_siec(sK0) <- (~13) [resolution 193,87] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
199. $false <- (3, ~13) [subsumption resolution 197,124] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
200. ~3 | 13 [avatar contradiction clause 199]
206. ~l1_e_siec(sK0) | ~v3_e_siec(sK0) | ~v2_e_siec(sK0) <- (~12) [resolution 185,93] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
208. ~v3_e_siec(sK0) | ~v2_e_siec(sK0) <- (3, ~12) [subsumption resolution 206,124] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
209. ~v2_e_siec(sK0) <- (3, 4, ~12) [subsumption resolution 208,129] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
210. $false <- (3, 4, 5, ~12) [subsumption resolution 209,134] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
211. ~3 | ~4 | ~5 | 12 [avatar contradiction clause 210]
219. ~r1_tarski(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0))),k2_zfmisc_1(k8_e_siec(sK0),k7_e_siec(sK0))) | ~v1_relat_1(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))) <- (~2) [resolution 119,97] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
221. 15 <=> v1_relat_1(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))) [avatar definition]
223. ~v1_relat_1(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0)))) <- (~15) [avatar component clause 221] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
225. 16 <=> r1_tarski(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0))),k2_zfmisc_1(k8_e_siec(sK0),k7_e_siec(sK0))) [avatar definition]
227. ~r1_tarski(k4_xboole_0(u2_e_siec(sK0),k4_relat_1(u1_struct_0(sK0))),k2_zfmisc_1(k8_e_siec(sK0),k7_e_siec(sK0))) <- (~16) [avatar component clause 225] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
228. ~15 | ~16 | 2 [avatar split clause 219,117,225,221]
229. ~v1_relat_1(u2_e_siec(sK0)) <- (~15) [resolution 223,101] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
231. 17 <=> v1_relat_1(u2_e_siec(sK0)) [avatar definition]
233. ~v1_relat_1(u2_e_siec(sK0)) <- (~17) [avatar component clause 231] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
234. ~17 | 15 [avatar split clause 229,221,231]
237. ~l1_e_siec(sK0) <- (~17) [resolution 233,86] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
239. $false <- (3, ~17) [subsumption resolution 237,124] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
240. ~3 | 17 [avatar contradiction clause 239]
242. ~l1_e_siec(sK0) | ~v3_e_siec(sK0) | ~v2_e_siec(sK0) <- (~16) [resolution 227,94] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
244. ~v3_e_siec(sK0) | ~v2_e_siec(sK0) <- (3, ~16) [subsumption resolution 242,124] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
245. ~v2_e_siec(sK0) <- (3, 4, ~16) [subsumption resolution 244,129] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
246. $false <- (3, 4, 5, ~16) [subsumption resolution 245,134] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
247. ~3 | ~4 | ~5 | 16 [avatar contradiction clause 246]
248. $false [avatar sat refutation 120,125,130,135,186,194,200,211,228,234,240,247] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:110,thDist:-110}
% SZS output end Proof for e_siec__t33_e_siec
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 36
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 105
% Active clauses: 50
% Passive clauses: 62
% Final active clauses: 44
% Final passive clauses: 11
% Main loop iterations started: 51
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 23
% Unit resulting resolution: 4
% Forward superposition: 3
% Backward superposition: 3
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 8
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

