i: [1,93,0,2,1,1,0,0] t69_ordinal3
a: [3,141,0,2,1,93]
s: 141
i: [1,92,0,2,1,1,0,0] t69_ordinal3
a: [3,146,0,2,1,92]
s: 146
i: [1,94,0,7,1,1,0,0] t69_ordinal3
a: [3,136,0,7,1,94]
s: 136
i: [1,95,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,151,0,2,1,95]
s: 151
i: [1,96,0,2,1,0,0,1] fc1_xboole_0
a: [3,156,0,2,1,96]
s: 156
i: [1,123,0,2,1,0,0,2] rc2_xboole_0
a: [3,166,0,2,1,123]
s: 166
i: [1,124,0,2,1,0,0,255] rc4_ordinal1
a: [3,171,0,2,1,124]
s: 171
i: [1,125,0,2,1,0,0,2] rc1_xboole_0
a: [3,176,0,2,1,125]
s: 176
i: [1,126,0,2,1,0,0,1] rc1_ordinal1
a: [3,181,0,2,1,126]
s: 181
i: [1,130,0,2,1,0,0,255] rc2_ordinal1
a: [3,186,0,2,1,130]
s: 186
i: [1,129,0,2,1,0,0,255] rc2_ordinal1
a: [3,191,0,2,1,129]
s: 191
i: [1,128,0,2,1,0,0,255] rc2_ordinal1
a: [3,196,0,2,1,128]
s: 196
i: [1,127,0,2,1,0,0,255] rc2_ordinal1
a: [3,201,0,2,1,127]
s: 201
i: [1,132,0,2,1,0,0,255] rc5_ordinal1
a: [3,206,0,2,1,132]
s: 206
i: [1,131,0,2,1,0,0,255] rc5_ordinal1
a: [3,211,0,2,1,131]
s: 211
i: [1,97,0,3,1,0,0,255] d2_xboole_0
a: [3,161,0,3,1,97]
s: 161
r: 151
i: [1,98,0,4,2,0,0,255] cc6_ordinal1
s: 98
i: [1,99,0,4,2,0,0,255] cc7_ordinal1
s: 99
d: [2,214,1,2,1,1,40,98,171]
a: [3,219,1,2,1,214]
s: 219
i: [1,100,0,4,2,0,0,2] cc3_ordinal1
s: 100
i: [1,104,0,4,2,0,0,255] cc1_ordinal1
s: 104
d: [2,250,1,2,1,1,40,104,146]
a: [3,258,1,2,1,250]
s: 258
i: [1,105,0,4,2,0,0,255] cc1_ordinal1
s: 105
d: [2,251,1,2,1,1,40,104,141]
a: [3,263,1,2,1,251]
s: 263
i: [1,117,0,4,1,0,0,255] existence_m1_subset_1
s: 117
i: [1,101,0,5,2,0,0,1] t6_boole
s: 101
d: [2,275,1,2,1,1,40,105,146]
a: [3,283,1,2,1,275]
s: 283
i: [1,116,0,5,2,0,0,255] t7_boole
s: 116
d: [2,276,1,2,1,1,40,105,141]
a: [3,288,1,2,1,276]
s: 288
i: [1,112,0,6,3,0,0,255] cc2_ordinal1
s: 112
i: [1,113,0,6,2,0,0,255] t1_subset
s: 113
d: [2,215,1,2,1,1,40,98,206]
a: [3,224,1,2,1,215]
s: 224
i: [1,118,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 118
d: [2,226,1,2,1,1,40,99,176]
a: [3,231,1,2,1,226]
s: 231
i: [1,102,0,7,2,0,0,1] t38_ordinal2
s: 102
i: [1,103,0,7,2,0,0,2] t54_ordinal3
s: 103
d: [2,227,1,2,1,1,40,99,156]
a: [3,236,1,2,1,227]
s: 236
i: [1,106,0,7,2,0,0,2] t56_ordinal3
s: 106
d: [2,238,1,2,1,1,40,100,176]
a: [3,243,1,2,1,238]
s: 243
i: [1,107,0,7,2,0,0,2] t56_ordinal3
s: 107
i: [1,110,0,7,3,0,0,255] cc2_ordinal2
s: 110
d: [2,239,1,2,1,1,40,100,156]
a: [3,248,1,2,1,239]
s: 248
i: [1,111,0,7,3,0,0,255] cc5_ordinal1
s: 111
d: [2,253,1,2,1,1,40,104,181]
a: [3,273,1,2,1,253]
s: 273
i: [1,115,0,7,3,0,0,1] t8_boole
s: 115
i: [1,114,0,8,3,0,0,255] t2_subset
s: 114
d: [2,278,1,2,1,1,40,105,181]
a: [3,298,1,2,1,278]
s: 298
i: [1,119,0,8,3,0,0,2] dt_k6_ordinal3
s: 119
d: [2,252,2,2,1,1,40,104,219]
a: [3,268,2,2,1,252]
s: 268
i: [1,120,0,8,3,0,0,1] dt_k7_ordinal3
s: 120
i: [1,121,0,8,3,0,0,1] dt_k11_ordinal2
s: 121
d: [2,277,2,2,1,1,40,105,219]
a: [3,293,2,2,1,277]
s: 293
i: [1,122,0,8,3,0,0,2] dt_k5_ordinal3
s: 122
d: [2,305,2,2,1,1,40,224,105]
a: [3,310,2,2,1,305]
s: 310
i: [1,109,0,14,4,0,0,2] t68_ordinal3
s: 109
i: [1,108,0,15,3,0,0,1] d7_ordinal3
s: 108
d: [2,306,2,2,1,1,40,224,104]
a: [3,315,2,2,1,306]
s: 315
d: [2,337,1,5,2,0,40,110,117]
s: 337
d: [2,324,2,2,1,1,40,243,105]
a: [3,329,2,2,1,324]
s: 329
d: [2,351,1,6,2,0,40,114,117]
s: 351
d: [2,300,1,7,3,0,44,101,101]
s: 300
d: [2,325,2,2,1,1,40,243,104]
a: [3,334,2,2,1,325]
s: 334
d: [2,353,1,8,3,0,40,119,105]
s: 353
d: [2,338,2,2,1,1,40,248,105]
a: [3,344,2,2,1,338]
s: 344
d: [2,354,1,8,3,0,40,119,104]
s: 354
d: [2,355,1,8,3,0,40,120,105]
s: 355
d: [2,339,2,2,1,1,40,248,104]
a: [3,349,2,2,1,339]
s: 349
d: [2,356,1,8,3,0,40,120,104]
s: 356
d: [2,395,3,4,2,1,44,344,101]
s: 395
d: [2,357,1,8,3,0,40,121,105]
s: 357
d: [2,358,1,8,3,0,40,121,104]
s: 358
d: [2,399,3,4,2,1,44,349,101]
s: 399
d: [2,362,1,8,3,0,40,122,105]
s: 362
d: [2,385,2,5,2,0,40,337,105]
s: 385
d: [2,363,1,8,3,0,40,122,104]
s: 363
d: [2,319,1,9,3,0,44,102,101]
s: 319
d: [2,427,2,9,3,1,44,136,319]
d: [2,437,2,7,2,2,21,427,146]
a: [3,441,2,2,1,437]
s: 441
d: [2,323,1,9,3,0,44,106,101]
s: 323
d: [2,386,2,5,2,0,40,337,104]
s: 386
d: [2,336,1,9,3,0,44,107,101]
s: 336
d: [2,375,1,15,4,0,44,108,102]
d: [2,382,1,13,3,1,21,375,248]
s: 382
d: [2,478,2,13,4,1,44,382,103]
d: [2,499,2,11,3,1,18,478]
d: [2,501,2,9,2,1,21,499,248]
a: [3,505,2,4,1,501]
s: 505
d: [2,523,3,0,0,2,52,248,248,505,119]
e: 523
d: [2,524,3,4,2,1,40,505,119]
d: [2,526,3,2,1,1,18,524]
d: [2,528,3,0,0,2,21,526,248]
e: 528
r: 505
d: [2,376,1,15,4,0,44,108,107]
d: [2,383,1,13,3,1,21,376,248]
s: 383
a: [3,504,2,4,1,501]
s: 504
d: [2,378,1,15,4,0,44,107,108]
d: [2,384,1,13,3,1,21,378,248]
s: 384
d: [2,374,1,26,6,0,44,108,109]
d: [2,380,1,24,5,0,18,374]
d: [2,381,1,20,4,0,21,380,121]
s: 381
d: [2,540,3,4,1,1,40,504,105]
a: [3,546,3,4,1,540]
s: 546
a: [3,509,2,5,1,501]
s: 509
d: [2,541,3,4,1,1,40,504,104]
a: [3,551,3,4,1,541]
s: 551
d: [2,388,2,6,2,0,40,351,118]
s: 388
d: [2,482,2,13,4,1,44,106,382]
d: [2,496,2,11,3,1,18,482]
s: 496
d: [2,542,3,6,2,1,44,504,101]
s: 542
d: [2,477,2,13,4,1,44,382,106]
d: [2,500,2,11,3,1,18,477]
s: 500
d: [2,575,4,6,2,1,44,546,101]
s: 575
d: [2,479,2,15,5,1,44,382,336]
d: [2,511,2,13,4,1,21,479,248]
d: [2,512,2,11,3,1,21,511,100]
s: 512
d: [2,480,2,15,5,1,44,382,323]
d: [2,498,2,13,4,1,18,480]
d: [2,513,2,11,3,2,21,498,156]
s: 513
d: [2,583,4,6,2,1,44,551,101]
s: 583
d: [2,484,2,15,5,1,44,336,382]
d: [2,518,2,13,4,1,21,484,248]
d: [2,519,2,11,3,1,21,518,100]
s: 519
d: [2,591,4,6,2,1,40,542,105]
s: 591
d: [2,485,2,15,5,1,44,323,382]
d: [2,494,2,13,4,1,18,485]
d: [2,520,2,11,3,2,21,494,156]
s: 520
d: [2,530,2,13,4,1,40,383,121]
d: [2,536,2,11,3,1,18,530]
s: 536
d: [2,592,4,6,2,1,40,542,104]
s: 592
d: [2,533,2,15,5,1,44,383,319]
d: [2,534,2,13,4,1,18,533]
d: [2,539,2,11,3,1,21,534,100]
s: 539
d: [2,576,3,7,2,1,44,509,101]
s: 576
d: [2,555,2,18,5,0,44,381,103]
d: [2,568,2,14,4,0,21,555,121]
s: 568
d: [2,532,2,13,4,1,44,383,102]
d: [2,535,2,11,3,1,18,532]
d: [2,537,2,9,2,1,21,535,248]
d: [2,619,3,10,4,1,44,136,568]
d: [2,628,3,7,3,1,19,619]
d: [2,630,3,5,2,2,21,628,141]
d: [2,631,3,3,1,3,21,630,146]
f: 631,537,528,437,239,97,96,95,92,93,94
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t69_ordinal3
% SZS output start Proof for ordinal3__t69_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => k7_ordinal3(k11_ordinal2(X0,X1),X1) = k1_xboole_0)) [input t69_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => k7_ordinal3(k11_ordinal2(X0,X1),X1) = k1_xboole_0)) [negated conjecture 1]
7. ! [X0] : (v1_xboole_0(X0) => v3_ordinal1(X0)) [input cc3_ordinal1]
13. k1_xboole_0 = o_0_0_xboole_0 [input d2_xboole_0]
14. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => k7_ordinal3(X0,X1) = k5_ordinal3(X0,k11_ordinal2(k6_ordinal3(X0,X1),X1)))) [input d7_ordinal3]
15. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k11_ordinal2(X0,X1))) [input dt_k11_ordinal2]
18. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k6_ordinal3(X0,X1))) [input dt_k6_ordinal3]
21. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
23. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
32. ! [X0] : (v3_ordinal1(X0) => k1_xboole_0 = k11_ordinal2(X0,k1_xboole_0)) [input t38_ordinal2]
33. ! [X0] : (v3_ordinal1(X0) => k1_xboole_0 = k5_ordinal3(X0,X0)) [input t54_ordinal3]
34. ! [X0] : (v3_ordinal1(X0) => (k1_xboole_0 = k5_ordinal3(k1_xboole_0,X0) & k5_ordinal3(X0,k1_xboole_0) = X0)) [input t56_ordinal3]
35. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (k1_xboole_0 != X0 => k6_ordinal3(k11_ordinal2(X1,X0),X0) = X1))) [input t68_ordinal3]
36. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
45. ? [X0] : (? [X1] : (k7_ordinal3(k11_ordinal2(X0,X1),X1) != k1_xboole_0 & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
48. ! [X0] : (v3_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 7]
49. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 36]
50. ! [X0] : (k1_xboole_0 = k11_ordinal2(X0,k1_xboole_0) | ~v3_ordinal1(X0)) [ennf transformation 32]
51. ! [X0] : (k1_xboole_0 = k5_ordinal3(X0,X0) | ~v3_ordinal1(X0)) [ennf transformation 33]
53. ! [X0] : ((k1_xboole_0 = k5_ordinal3(k1_xboole_0,X0) & k5_ordinal3(X0,k1_xboole_0) = X0) | ~v3_ordinal1(X0)) [ennf transformation 34]
54. ! [X0] : (! [X1] : (k7_ordinal3(X0,X1) = k5_ordinal3(X0,k11_ordinal2(k6_ordinal3(X0,X1),X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 14]
55. ! [X0] : (! [X1] : ((k6_ordinal3(k11_ordinal2(X1,X0),X0) = X1 | k1_xboole_0 = X0) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 35]
56. ! [X0] : (! [X1] : (k6_ordinal3(k11_ordinal2(X1,X0),X0) = X1 | k1_xboole_0 = X0 | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 55]
67. ! [X0,X1] : (v3_ordinal1(k6_ordinal3(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 18]
68. ! [X0,X1] : (v3_ordinal1(k6_ordinal3(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 67]
71. ! [X0,X1] : (v3_ordinal1(k11_ordinal2(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 15]
72. ! [X0,X1] : (v3_ordinal1(k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 71]
75. ? [X0] : (? [X1] : (k7_ordinal3(k11_ordinal2(X0,X1),X1) != k1_xboole_0 & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (k1_xboole_0 != k7_ordinal3(k11_ordinal2(sK0,X1),X1) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
76. ? [X1] : (k1_xboole_0 != k7_ordinal3(k11_ordinal2(sK0,X1),X1) & v3_ordinal1(X1)) => (k1_xboole_0 != k7_ordinal3(k11_ordinal2(sK0,sK1),sK1) & v3_ordinal1(sK1)) [choice axiom]
77. (k1_xboole_0 != k7_ordinal3(k11_ordinal2(sK0,sK1),sK1) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 45,76,75]
92. v3_ordinal1(sK0) [cnf transformation 77] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
93. v3_ordinal1(sK1) [cnf transformation 77] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
94. k1_xboole_0 != k7_ordinal3(k11_ordinal2(sK0,sK1),sK1) [cnf transformation 77] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
95. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 21] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
96. v1_xboole_0(k1_xboole_0) [cnf transformation 23] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
97. k1_xboole_0 = o_0_0_xboole_0 [cnf transformation 13] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
100. ~v1_xboole_0(X0) | v3_ordinal1(X0) [cnf transformation 48] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
101. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 49] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
102. k1_xboole_0 = k11_ordinal2(X0,k1_xboole_0) | ~v3_ordinal1(X0) [cnf transformation 50] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
103. k1_xboole_0 = k5_ordinal3(X0,X0) | ~v3_ordinal1(X0) [cnf transformation 51] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
107. k1_xboole_0 = k5_ordinal3(k1_xboole_0,X0) | ~v3_ordinal1(X0) [cnf transformation 53] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
108. k7_ordinal3(X0,X1) = k5_ordinal3(X0,k11_ordinal2(k6_ordinal3(X0,X1),X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 54] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
109. k6_ordinal3(k11_ordinal2(X1,X0),X0) = X1 | k1_xboole_0 = X0 | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 56] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
119. v3_ordinal1(k6_ordinal3(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 68] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
121. v3_ordinal1(k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 72] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
134. 1 <=> k1_xboole_0 = k7_ordinal3(k11_ordinal2(sK0,sK1),sK1) [avatar definition]
136. k1_xboole_0 != k7_ordinal3(k11_ordinal2(sK0,sK1),sK1) <- (~1) [avatar component clause 134] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
137. ~1 [avatar split clause 94,134]
139. 2 <=> v3_ordinal1(sK1) [avatar definition]
141. v3_ordinal1(sK1) <- (2) [avatar component clause 139] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. 2 [avatar split clause 93,139]
144. 3 <=> v3_ordinal1(sK0) [avatar definition]
146. v3_ordinal1(sK0) <- (3) [avatar component clause 144] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
147. 3 [avatar split clause 92,144]
149. 4 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
152. 4 [avatar split clause 95,149]
154. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
156. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 154] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
157. 5 [avatar split clause 96,154]
159. 6 <=> k1_xboole_0 = o_0_0_xboole_0 [avatar definition]
162. 6 [avatar split clause 97,159]
239. v3_ordinal1(k1_xboole_0) <- (5) [resolution 100,156] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
246. 22 <=> v3_ordinal1(k1_xboole_0) [avatar definition]
248. v3_ordinal1(k1_xboole_0) <- (22) [avatar component clause 246] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
249. 22 | ~5 [avatar split clause 239,154,246]
319. k11_ordinal2(X1,X0) = X0 | ~v3_ordinal1(X1) | ~v1_xboole_0(X0) [superposition 102,101] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
374. k7_ordinal3(k11_ordinal2(X0,X1),X1) = k5_ordinal3(k11_ordinal2(X0,X1),k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k11_ordinal2(X0,X1)) | k1_xboole_0 = X1 | ~v3_ordinal1(X0) | ~v3_ordinal1(X1) [superposition 108,109] {a:1,w:26,wCS:10400,sine:1,thAx:0,allAx:2,thDist:-2}
376. k1_xboole_0 = k7_ordinal3(k1_xboole_0,X2) | ~v3_ordinal1(X2) | ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(k11_ordinal2(k6_ordinal3(k1_xboole_0,X2),X2)) [superposition 108,107] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
380. k7_ordinal3(k11_ordinal2(X0,X1),X1) = k5_ordinal3(k11_ordinal2(X0,X1),k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k11_ordinal2(X0,X1)) | k1_xboole_0 = X1 | ~v3_ordinal1(X0) [duplicate literal removal 374] {a:1,w:24,wCS:9600,sine:1,thAx:0,allAx:2,thDist:-2}
381. k7_ordinal3(k11_ordinal2(X0,X1),X1) = k5_ordinal3(k11_ordinal2(X0,X1),k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | k1_xboole_0 = X1 | ~v3_ordinal1(X0) [subsumption resolution 380,121] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
383. ~v3_ordinal1(k11_ordinal2(k6_ordinal3(k1_xboole_0,X2),X2)) | ~v3_ordinal1(X2) | k1_xboole_0 = k7_ordinal3(k1_xboole_0,X2) <- (22) [subsumption resolution 376,248] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
427. k1_xboole_0 != k7_ordinal3(sK1,sK1) | ~v3_ordinal1(sK0) | ~v1_xboole_0(sK1) <- (~1) [superposition 136,319] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
437. k1_xboole_0 != k7_ordinal3(sK1,sK1) | ~v1_xboole_0(sK1) <- (~1, 3) [subsumption resolution 427,146] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
439. 37 <=> v1_xboole_0(sK1) [avatar definition]
443. 38 <=> k1_xboole_0 = k7_ordinal3(sK1,sK1) [avatar definition]
446. ~37 | ~38 | 1 | ~3 [avatar split clause 437,144,134,443,439]
503. 39 <=> v3_ordinal1(k6_ordinal3(k1_xboole_0,k1_xboole_0)) [avatar definition]
505. ~v3_ordinal1(k6_ordinal3(k1_xboole_0,k1_xboole_0)) <- (~39) [avatar component clause 503] {a:2,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
507. 40 <=> k1_xboole_0 = k7_ordinal3(k1_xboole_0,k1_xboole_0) [avatar definition]
524. ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(k1_xboole_0) <- (~39) [resolution 505,119] {a:3,w:4,wCS:1600,sine:2,thAx:0,allAx:4,thDist:-4}
526. ~v3_ordinal1(k1_xboole_0) <- (~39) [duplicate literal removal 524] {a:3,w:2,wCS:800,sine:2,thAx:0,allAx:4,thDist:-4}
528. $false <- (22, ~39) [subsumption resolution 526,248] {a:3,w:0,sine:2,thAx:0,allAx:4,thDist:-4}
529. ~22 | 39 [avatar contradiction clause 528]
532. ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(k1_xboole_0) | k1_xboole_0 = k7_ordinal3(k1_xboole_0,k1_xboole_0) | ~v3_ordinal1(k6_ordinal3(k1_xboole_0,k1_xboole_0)) <- (22) [superposition 383,102] {a:2,w:13,wCS:5200,sine:1,thAx:0,allAx:3,thDist:-3}
535. ~v3_ordinal1(k1_xboole_0) | k1_xboole_0 = k7_ordinal3(k1_xboole_0,k1_xboole_0) | ~v3_ordinal1(k6_ordinal3(k1_xboole_0,k1_xboole_0)) <- (22) [duplicate literal removal 532] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:3,thDist:-3}
537. k1_xboole_0 = k7_ordinal3(k1_xboole_0,k1_xboole_0) | ~v3_ordinal1(k6_ordinal3(k1_xboole_0,k1_xboole_0)) <- (22) [subsumption resolution 535,248] {a:2,w:9,wCS:3600,sine:1,thAx:0,allAx:3,thDist:-3}
538. ~39 | 40 | ~22 [avatar split clause 537,246,507,503]
555. k1_xboole_0 = k7_ordinal3(k11_ordinal2(X2,X3),X3) | ~v3_ordinal1(X3) | k1_xboole_0 = X3 | ~v3_ordinal1(X2) | ~v3_ordinal1(k11_ordinal2(X2,X3)) [superposition 381,103] {a:2,w:18,wCS:7200,sine:1,thAx:0,allAx:3,thDist:-3}
568. k1_xboole_0 = k7_ordinal3(k11_ordinal2(X2,X3),X3) | ~v3_ordinal1(X3) | k1_xboole_0 = X3 | ~v3_ordinal1(X2) [subsumption resolution 555,121] {a:2,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
619. k1_xboole_0 != k1_xboole_0 | ~v3_ordinal1(sK1) | k1_xboole_0 = sK1 | ~v3_ordinal1(sK0) <- (~1) [superposition 136,568] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
628. ~v3_ordinal1(sK1) | k1_xboole_0 = sK1 | ~v3_ordinal1(sK0) <- (~1) [trivial inequality removal 619] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
630. k1_xboole_0 = sK1 | ~v3_ordinal1(sK0) <- (~1, 2) [subsumption resolution 628,141] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
631. k1_xboole_0 = sK1 <- (~1, 2, 3) [subsumption resolution 630,146] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
633. 43 <=> k1_xboole_0 = sK1 [avatar definition]
636. 43 | 1 | ~2 | ~3 [avatar split clause 631,144,139,134,633]
637. k1_xboole_0 != sK1 | k1_xboole_0 != o_0_0_xboole_0 | v1_xboole_0(sK1) | ~v1_xboole_0(o_0_0_xboole_0) [theory tautology sat conflict] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:1,thDist:-1}
638. k1_xboole_0 != sK1 | k1_xboole_0 != k7_ordinal3(k1_xboole_0,k1_xboole_0) | k1_xboole_0 = k7_ordinal3(sK1,sK1) [theory tautology sat conflict] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:1,thDist:-1}
639. $false [avatar sat refutation 137,142,147,152,157,162,249,446,529,538,636,637,638] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:242,thDist:-242}
% SZS output end Proof for ordinal3__t69_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 37
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 359
% Active clauses: 96
% Passive clauses: 121
% Final active clauses: 94
% Final passive clauses: 16
% Main loop iterations started: 105
% 
% >>> Simplifying Inferences
% Duplicate literals: 73
% Trivial inequalities: 1
% Fw subsumption resolutions: 29
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 82
% 
% >>> Generating Inferences
% Binary resolution: 64
% Unit resulting resolution: 1
% Forward superposition: 53
% Backward superposition: 56
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 53
% SAT solver unit clauses: 16
% SAT solver binary clauses: 26
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.002 s
% other: 0.004 s ( own 0.002 s ) 

