i: [1,85,0,3,1,1,0,0] l17_arytm_1
a: [3,133,0,3,1,85]
s: 133
i: [1,84,0,3,1,1,0,0] l17_arytm_1
a: [3,138,0,3,1,84]
s: 138
i: [1,83,0,3,1,1,0,0] l17_arytm_1
a: [3,143,0,3,1,83]
s: 143
i: [1,86,0,5,1,1,0,0] l17_arytm_1
a: [3,128,0,5,1,86]
s: 128
d: [2,183,0,5,1,2,22,128,133]
a: [3,187,0,5,1,183]
s: 187
i: [1,106,0,3,1,0,0,255] reflexivity_r1_tarski
s: 106
i: [1,87,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,148,0,3,1,87]
s: 148
i: [1,88,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,153,0,3,1,88]
s: 153
i: [1,118,0,11,3,0,0,255] d8_arytm_2
a: [3,162,0,3,1,118]
s: 162
i: [1,105,0,4,1,0,0,1] existence_m1_subset_1
s: 105
i: [1,89,0,5,1,0,0,255] dt_k1_arytm_2
a: [3,158,0,5,1,89]
s: 158
i: [1,100,0,6,2,0,0,255] t1_subset
s: 100
i: [1,107,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 107
i: [1,120,0,9,3,0,0,2] t6_arytm_1
a: [3,174,0,6,2,120]
s: 174
i: [1,90,0,7,2,0,0,255] dt_k4_arytm_2
s: 90
i: [1,102,0,7,2,0,0,255] t3_subset
s: 102
d: [2,190,1,3,1,2,40,174,138]
a: [3,200,1,3,1,190]
s: 200
i: [1,101,0,8,3,0,0,255] t2_subset
s: 101
d: [2,209,1,5,2,1,40,101,153]
a: [3,222,1,2,1,209]
s: 222
a: [3,166,0,8,2,118]
s: 166
i: [1,119,0,11,3,0,0,255] d8_arytm_2
a: [3,170,0,8,2,119]
s: 170
d: [2,210,1,5,2,1,40,101,162]
a: [3,231,1,2,1,210]
s: 231
i: [1,91,0,9,2,0,0,255] dt_k3_arytm_2
s: 91
d: [2,189,1,3,1,2,40,174,162]
a: [3,195,1,3,1,189]
s: 195
i: [1,104,0,9,3,0,0,255] t5_subset
s: 104
i: [1,103,0,10,3,0,0,255] t4_subset
s: 103
d: [2,252,1,6,2,1,40,104,158]
a: [3,260,1,3,1,252]
s: 260
i: [1,114,0,11,3,0,0,1] dt_k1_arytm_1
s: 114
d: [2,191,1,4,1,1,40,174,105]
a: [3,205,1,4,1,191]
s: 205
i: [1,115,0,11,3,0,0,2] dt_k7_arytm_2
s: 115
i: [1,117,0,12,4,0,0,2] connectedness_r1_arytm_2
s: 117
d: [2,277,1,9,3,1,40,117,138]
s: 277
i: [1,112,0,13,3,0,0,255] commutativity_k5_arytm_2
s: 112
d: [2,283,2,6,2,1,40,277,138]
d: [2,285,2,3,1,1,18,283]
a: [3,289,2,3,1,285]
s: 289
i: [1,113,0,13,3,0,0,255] dt_k5_arytm_2
s: 113
i: [1,116,0,13,3,0,0,2] commutativity_k7_arytm_2
s: 116
d: [2,284,2,8,2,1,40,277,105]
a: [3,294,2,4,1,284]
s: 294
i: [1,111,0,14,4,0,0,255] existence_m2_subset_1
s: 111
d: [2,282,2,11,3,1,40,277,90]
s: 282
i: [1,121,0,17,5,0,0,1] d1_arytm_1
a: [3,178,0,14,4,121]
s: 178
i: [1,108,0,15,5,0,0,255] dt_m2_subset_1
s: 108
r: 260
a: [3,256,1,3,1,252]
s: 256
i: [1,109,0,15,5,0,0,255] redefinition_m2_subset_1
s: 109
a: [3,259,1,3,1,252]
s: 259
i: [1,110,0,15,5,0,0,255] redefinition_m2_subset_1
s: 110
i: [1,122,0,19,5,0,0,1] d1_arytm_1
s: 122
d: [2,280,2,16,4,1,40,277,114]
s: 280
i: [1,124,0,21,5,0,0,1] d1_arytm_1
s: 124
d: [2,281,2,16,4,1,40,277,115]
s: 281
d: [2,368,0,16,4,0,21,124,114]
s: 368
i: [1,94,0,22,5,0,0,255] d8_arytm_2
s: 94
d: [2,211,1,9,2,1,40,101,158]
a: [3,236,1,5,1,211]
s: 236
i: [1,123,0,23,5,0,0,1] d1_arytm_1
s: 123
d: [2,214,1,6,2,0,40,101,105]
s: 214
r: 256
r: 259
d: [2,407,0,18,4,0,21,123,115]
s: 407
d: [2,327,1,18,5,0,40,108,91]
d: [2,331,1,13,4,1,21,327,158]
d: [2,332,1,10,3,2,21,331,260]
a: [3,336,1,2,1,332]
s: 336
s: 260
d: [2,207,1,7,2,1,40,90,174]
s: 207
d: [2,405,2,5,1,1,40,236,107]
a: [3,415,2,5,1,405]
s: 415
d: [2,253,1,7,2,0,40,104,105]
s: 253
d: [2,263,1,7,2,1,40,103,158]
s: 263
d: [2,408,2,6,2,0,40,214,107]
s: 408
d: [2,251,1,8,3,0,40,104,102]
s: 251
d: [2,449,2,5,2,0,40,251,106]
s: 449
a: [3,239,1,4,1,211]
s: 239
d: [2,264,1,8,2,0,40,103,105]
s: 264
d: [2,442,2,6,2,0,40,253,214]
s: 442
d: [2,421,1,17,5,1,44,407,166]
d: [2,436,1,14,4,1,18,421]
d: [2,437,1,11,3,2,21,436,162]
d: [2,438,1,8,2,2,21,437,117]
s: 438
d: [2,445,2,8,3,1,40,263,104]
d: [2,447,2,6,2,2,21,445,222]
s: 447
d: [2,422,1,17,5,1,44,407,170]
d: [2,435,1,14,4,1,18,422]
d: [2,439,1,11,3,2,21,435,162]
d: [2,440,1,8,2,3,21,439,174]
s: 440
d: [2,462,2,0,0,5,52,138,187,440]
e: 462
d: [2,465,2,6,2,4,44,187,440]
d: [2,475,2,3,1,4,19,465]
d: [2,478,2,0,0,5,21,475,138]
e: 478
d: [2,181,0,3,1,2,22,143,133]
f: 478,183,181,120,119,83,84,85,86
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__l17_arytm_1
% SZS output start Proof for arytm_1__l17_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => k1_arytm_1(X1,X0) = X1))) [input l17_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => k1_arytm_1(X1,X0) = X1))) [negated conjecture 1]
7. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => ((~r1_arytm_2(X1,X0) => (k1_arytm_1(X0,X1) = X2 <=> k11_arytm_3 = X2)) & (r1_arytm_2(X1,X0) => (k1_arytm_1(X0,X1) = X2 <=> k7_arytm_2(X2,X1) = X0)))))) [input d1_arytm_1]
8. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (~(k7_arytm_2(X0,X1) != k4_arytm_2(k5_arytm_2(k3_arytm_2(X0),k3_arytm_2(X1))) & k11_arytm_3 != X0 & k11_arytm_3 != X1) & (k11_arytm_3 = X0 => k7_arytm_2(X0,X1) = X1) & (k11_arytm_3 = X1 => k7_arytm_2(X0,X1) = X0)))) [input d8_arytm_2]
19. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2)) [input dt_k7_arytm_2]
32. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => r1_arytm_2(X0,X1)))) [input t6_arytm_1]
40. ? [X0] : (? [X1] : ((k1_arytm_1(X1,X0) != X1 & k11_arytm_3 = X0) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
41. ? [X0] : (? [X1] : (k1_arytm_1(X1,X0) != X1 & k11_arytm_3 = X0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 40]
44. ! [X0] : (! [X1] : (((k7_arytm_2(X0,X1) = k4_arytm_2(k5_arytm_2(k3_arytm_2(X0),k3_arytm_2(X1))) | k11_arytm_3 = X0 | k11_arytm_3 = X1) & (k7_arytm_2(X0,X1) = X1 | k11_arytm_3 != X0) & (k7_arytm_2(X0,X1) = X0 | k11_arytm_3 != X1)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 8]
45. ! [X0] : (! [X1] : ((r1_arytm_2(X0,X1) | k11_arytm_3 != X0) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 32]
46. ! [X0] : (! [X1] : (r1_arytm_2(X0,X1) | k11_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 45]
47. ! [X0] : (! [X1] : (! [X2] : ((((k1_arytm_1(X0,X1) = X2 <=> k11_arytm_3 = X2) | r1_arytm_2(X1,X0)) & ((k1_arytm_1(X0,X1) = X2 <=> k7_arytm_2(X2,X1) = X0) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 7]
68. ! [X0,X1] : (m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 19]
69. ! [X0,X1] : (m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 68]
74. ? [X0] : (? [X1] : (k1_arytm_1(X1,X0) != X1 & k11_arytm_3 = X0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (k1_arytm_1(X1,sK0) != X1 & k11_arytm_3 = sK0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
75. ? [X1] : (k1_arytm_1(X1,sK0) != X1 & k11_arytm_3 = sK0 & m1_subset_1(X1,k2_arytm_2)) => (sK1 != k1_arytm_1(sK1,sK0) & k11_arytm_3 = sK0 & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
76. (sK1 != k1_arytm_1(sK1,sK0) & k11_arytm_3 = sK0 & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 41,75,74]
77. ! [X0] : (! [X1] : (! [X2] : (((((k1_arytm_1(X0,X1) = X2 | k11_arytm_3 != X2) & (k11_arytm_3 = X2 | k1_arytm_1(X0,X1) != X2)) | r1_arytm_2(X1,X0)) & (((k1_arytm_1(X0,X1) = X2 | k7_arytm_2(X2,X1) != X0) & (k7_arytm_2(X2,X1) = X0 | k1_arytm_1(X0,X1) != X2)) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [nnf transformation 47]
83. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 76] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
84. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 76] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
85. k11_arytm_3 = sK0 [cnf transformation 76] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
86. sK1 != k1_arytm_1(sK1,sK0) [cnf transformation 76] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
92. k7_arytm_2(X0,X1) = X0 | k11_arytm_3 != X1 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 44] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
95. r1_arytm_2(X0,X1) | k11_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 46] {a:0,w:12,wCS:4800,sine:2,thAx:0,allAx:1,thDist:-1}
97. k1_arytm_1(X0,X1) = X2 | k7_arytm_2(X2,X1) != X0 | ~r1_arytm_2(X1,X0) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 77] {a:0,w:22,wCS:8800,sine:1,thAx:0,allAx:1,thDist:-1}
115. m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 69] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
119. k7_arytm_2(X0,k11_arytm_3) = X0 | ~m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [equality resolution 92] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
120. r1_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) [equality resolution 95] {a:0,w:9,wCS:3600,sine:2,thAx:0,allAx:1,thDist:-1}
123. k1_arytm_1(k7_arytm_2(X2,X1),X1) = X2 | ~r1_arytm_2(X1,k7_arytm_2(X2,X1)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k7_arytm_2(X2,X1),k2_arytm_2) [equality resolution 97] {a:0,w:23,wCS:9200,sine:1,thAx:0,allAx:1,thDist:-1}
126. 1 <=> sK1 = k1_arytm_1(sK1,sK0) [avatar definition]
128. sK1 != k1_arytm_1(sK1,sK0) <- (~1) [avatar component clause 126] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. ~1 [avatar split clause 86,126]
131. 2 <=> k11_arytm_3 = sK0 [avatar definition]
133. k11_arytm_3 = sK0 <- (2) [avatar component clause 131] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. 2 [avatar split clause 85,131]
136. 3 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
138. m1_subset_1(sK1,k2_arytm_2) <- (3) [avatar component clause 136] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
139. 3 [avatar split clause 84,136]
141. 4 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
143. m1_subset_1(sK0,k2_arytm_2) <- (4) [avatar component clause 141] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. 4 [avatar split clause 83,141]
161. 8 <=> m1_subset_1(k11_arytm_3,k2_arytm_2) [avatar definition]
162. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (8) [avatar component clause 161] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
169. 10 <=> ! [X0] : (k7_arytm_2(X0,k11_arytm_3) = X0 | ~m1_subset_1(X0,k2_arytm_2)) [avatar definition]
170. k7_arytm_2(X0,k11_arytm_3) = X0 | ~m1_subset_1(X0,k2_arytm_2) <- (10) [avatar component clause 169] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
171. ~8 | 10 [avatar split clause 119,169,161]
173. 11 <=> ! [X1] : (r1_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2)) [avatar definition]
174. ~m1_subset_1(X1,k2_arytm_2) | r1_arytm_2(k11_arytm_3,X1) <- (11) [avatar component clause 173] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
175. ~8 | 11 [avatar split clause 120,173,161]
181. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (2, 4) [forward demodulation 143,133] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
182. 8 | ~2 | ~4 [avatar split clause 181,141,131,161]
183. sK1 != k1_arytm_1(sK1,k11_arytm_3) <- (~1, 2) [forward demodulation 128,133] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
185. 13 <=> sK1 = k1_arytm_1(sK1,k11_arytm_3) [avatar definition]
187. sK1 != k1_arytm_1(sK1,k11_arytm_3) <- (~13) [avatar component clause 185] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
188. ~13 | 1 | ~2 [avatar split clause 183,131,126,185]
407. k1_arytm_1(k7_arytm_2(X2,X1),X1) = X2 | ~r1_arytm_2(X1,k7_arytm_2(X2,X1)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) [subsumption resolution 123,115] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
422. k1_arytm_1(X1,k11_arytm_3) = X1 | ~r1_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) <- (10) [superposition 407,170] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
435. k1_arytm_1(X1,k11_arytm_3) = X1 | ~r1_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) <- (10) [duplicate literal removal 422] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
439. k1_arytm_1(X1,k11_arytm_3) = X1 | ~r1_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2) <- (8, 10) [subsumption resolution 435,162] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
440. k1_arytm_1(X1,k11_arytm_3) = X1 | ~m1_subset_1(X1,k2_arytm_2) <- (8, 10, 11) [subsumption resolution 439,174] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
465. sK1 != sK1 | ~m1_subset_1(sK1,k2_arytm_2) <- (8, 10, 11, ~13) [superposition 187,440] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
475. ~m1_subset_1(sK1,k2_arytm_2) <- (8, 10, 11, ~13) [trivial inequality removal 465] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
478. $false <- (3, 8, 10, 11, ~13) [subsumption resolution 475,138] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
479. ~3 | ~8 | ~10 | ~11 | 13 [avatar contradiction clause 478]
480. $false [avatar sat refutation 129,134,139,144,171,175,182,188,479] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:187,thDist:-187}
% SZS output end Proof for arytm_1__l17_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 31
% >>> Preprocessing
% Introduced skolems: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 35
% Generated clauses: 252
% Active clauses: 65
% Passive clauses: 97
% Final active clauses: 62
% Final passive clauses: 27
% Main loop iterations started: 70
% 
% >>> Simplifying Inferences
% Duplicate literals: 64
% Trivial inequalities: 1
% Fw subsumption resolutions: 23
% Fw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 43
% 
% >>> Generating Inferences
% Binary resolution: 58
% Unit resulting resolution: 1
% Forward superposition: 27
% Backward superposition: 27
% 
% >>> AVATAR
% Split clauses: 18
% Split components: 41
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 39
% SAT solver unit clauses: 7
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.003 s

