i: [1,62,0,2,1,1,0,0] t15_yellow_5
a: [3,99,0,2,1,62]
s: 99
i: [1,61,0,2,1,1,0,0] t15_yellow_5
a: [3,104,0,2,1,61]
s: 104
i: [1,60,0,2,1,1,0,0] t15_yellow_5
a: [3,109,0,2,1,60]
s: 109
i: [1,59,0,2,1,1,0,0] t15_yellow_5
a: [3,114,0,2,1,59]
s: 114
i: [1,58,0,2,1,1,0,0] t15_yellow_5
a: [3,119,0,2,1,58]
s: 119
i: [1,65,0,8,2,1,0,0] t15_yellow_5
a: [3,84,0,4,1,65]
s: 84
a: [3,87,0,4,1,65]
s: 87
i: [1,63,0,4,1,1,0,0] t15_yellow_5
a: [3,94,0,4,1,63]
s: 94
i: [1,80,0,2,1,0,0,1] existence_l1_orders_2
a: [3,124,0,2,1,80]
s: 124
i: [1,72,0,4,1,0,0,1] existence_m1_subset_1
s: 72
i: [1,67,0,6,3,0,0,1] cc1_lattice3
s: 67
i: [1,68,0,6,3,0,0,2] cc2_lattice3
s: 68
i: [1,66,0,7,2,0,0,1] dt_k3_yellow_0
s: 66
i: [1,79,0,17,4,0,0,255] dt_k11_lattice3
s: 79
i: [1,71,0,18,6,0,0,1] t2_yellow_5
s: 71
i: [1,73,0,21,6,0,0,2] dt_k12_lattice3
s: 73
i: [1,76,0,22,7,0,0,1] symmetry_r2_yellow_5
s: 76
i: [1,77,0,22,7,0,0,255] redefinition_r2_yellow_5
s: 77
i: [1,78,0,22,7,0,0,255] redefinition_r2_yellow_5
s: 78
i: [1,69,0,23,6,0,0,255] d3_yellow_5
s: 69
i: [1,70,0,23,6,0,0,255] d3_yellow_5
s: 70
i: [1,74,0,23,6,0,0,255] redefinition_k12_lattice3
s: 74
i: [1,75,0,23,6,0,0,2] commutativity_k12_lattice3
s: 75
d: [2,132,1,37,11,0,44,74,70]
d: [2,139,1,27,8,0,18,132]
d: [2,140,1,25,7,0,21,139,68]
s: 140
d: [2,133,1,37,11,0,44,70,74]
d: [2,138,1,27,8,0,18,133]
d: [2,141,1,25,7,0,21,138,68]
s: 141
d: [2,134,1,37,11,0,44,69,74]
d: [2,137,1,27,8,0,18,134]
d: [2,142,1,25,7,0,21,137,68]
s: 142
d: [2,150,2,34,12,0,44,71,140]
d: [2,157,2,20,7,0,18,150]
s: 157
d: [2,147,2,34,12,0,44,140,71]
d: [2,160,2,20,7,0,18,147]
s: 160
d: [2,172,3,16,6,1,40,157,94]
d: [2,177,3,12,5,2,21,172,87]
d: [2,178,3,10,4,3,21,177,99]
d: [2,179,3,8,3,4,21,178,104]
d: [2,180,3,6,2,5,21,179,114]
d: [2,181,3,4,1,6,21,180,119]
a: [3,185,3,4,1,181]
s: 185
d: [2,187,4,0,0,6,52,114,104,99,94,94,84,185,78]
e: 187
d: [2,188,4,18,6,1,40,185,78]
d: [2,189,4,14,5,1,18,188]
d: [2,191,4,10,4,2,21,189,84]
d: [2,192,4,6,3,3,21,191,94]
d: [2,193,4,4,2,4,21,192,99]
d: [2,194,4,2,1,5,21,193,104]
d: [2,195,4,0,0,6,21,194,114]
e: 195
r: 84
r: 87
a: [3,83,0,4,1,65]
s: 83
a: [3,88,0,4,1,65]
s: 88
d: [2,161,2,34,12,0,44,142,71]
d: [2,168,2,20,7,0,18,161]
s: 168
d: [2,207,4,0,0,7,52,119,114,104,99,94,185,88,168]
e: 207
r: 185
i: [1,64,0,8,2,1,0,0] t15_yellow_5
d: [2,197,4,10,4,2,21,189,94]
d: [2,198,4,8,3,3,21,197,99]
d: [2,199,4,6,2,4,21,198,104]
d: [2,200,4,4,1,5,21,199,114]
d: [2,202,1,18,6,1,40,83,77]
d: [2,205,1,14,5,1,18,202]
d: [2,216,1,10,4,2,21,205,94]
d: [2,217,1,8,3,3,21,216,99]
d: [2,218,1,6,2,4,21,217,104]
d: [2,219,1,4,1,5,21,218,114]
f: 219,207,200,181,58,59,61,62,63,64,65
% Refutation found. Thanks to Tanya!
% SZS status Theorem for yellow_5__t15_yellow_5
% SZS output start Proof for yellow_5__t15_yellow_5
1. ! [X0] : ((l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => (r2_yellow_5(X0,X1,X1) <=> k3_yellow_0(X0) != X1))) [input t15_yellow_5]
2. ~! [X0] : ((l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => (r2_yellow_5(X0,X1,X1) <=> k3_yellow_0(X0) != X1))) [negated conjecture 1]
4. ! [X0] : (l1_orders_2(X0) => (v2_lattice3(X0) => ~v2_struct_0(X0))) [input cc2_lattice3]
6. ! [X0] : ((l1_orders_2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r1_yellow_5(X0,X1,X2) <=> k3_yellow_0(X0) != k11_lattice3(X0,X1,X2))))) [input d3_yellow_5]
17. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_orders_2(X0) & v2_lattice3(X0) & v5_orders_2(X0)) => k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2)) [input redefinition_k12_lattice3]
18. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_orders_2(X0) & v2_lattice3(X0) & v5_orders_2(X0)) => (r2_yellow_5(X0,X1,X2) <=> r1_yellow_5(X0,X1,X2))) [input redefinition_r2_yellow_5]
20. ! [X0] : ((l1_orders_2(X0) & v2_lattice3(X0) & v5_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k12_lattice3(X0,X1,X1) = X1)) [input t2_yellow_5]
23. ~! [X0] : ((l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => (r2_yellow_5(X0,X1,X1) <=> k3_yellow_0(X0) != X1))) [pure predicate removal 2]
24. ? [X0] : (? [X1] : ((r2_yellow_5(X0,X1,X1) <~> k3_yellow_0(X0) != X1) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v3_orders_2(X0))) [ennf transformation 23]
25. ? [X0] : (? [X1] : ((r2_yellow_5(X0,X1,X1) <~> k3_yellow_0(X0) != X1) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v3_orders_2(X0)) [flattening 24]
29. ! [X0] : ((~v2_struct_0(X0) | ~v2_lattice3(X0)) | ~l1_orders_2(X0)) [ennf transformation 4]
30. ! [X0] : (~v2_struct_0(X0) | ~v2_lattice3(X0) | ~l1_orders_2(X0)) [flattening 29]
31. ! [X0] : (! [X1] : (! [X2] : ((r1_yellow_5(X0,X1,X2) <=> k3_yellow_0(X0) != k11_lattice3(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 6]
32. ! [X0] : (! [X1] : (! [X2] : ((r1_yellow_5(X0,X1,X2) <=> k3_yellow_0(X0) != k11_lattice3(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [flattening 31]
33. ! [X0] : (! [X1] : (k12_lattice3(X0,X1,X1) = X1 | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v3_orders_2(X0))) [ennf transformation 20]
34. ! [X0] : (! [X1] : (k12_lattice3(X0,X1,X1) = X1 | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v3_orders_2(X0)) [flattening 33]
37. ! [X0,X1,X2] : (k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0))) [ennf transformation 17]
38. ! [X0,X1,X2] : (k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0)) [flattening 37]
43. ! [X0,X1,X2] : ((r2_yellow_5(X0,X1,X2) <=> r1_yellow_5(X0,X1,X2)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0))) [ennf transformation 18]
44. ! [X0,X1,X2] : ((r2_yellow_5(X0,X1,X2) <=> r1_yellow_5(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0)) [flattening 43]
47. ? [X0] : (? [X1] : (((k3_yellow_0(X0) = X1 | ~r2_yellow_5(X0,X1,X1)) & (k3_yellow_0(X0) != X1 | r2_yellow_5(X0,X1,X1))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v3_orders_2(X0)) [nnf transformation 25]
48. ? [X0] : (? [X1] : ((k3_yellow_0(X0) = X1 | ~r2_yellow_5(X0,X1,X1)) & (k3_yellow_0(X0) != X1 | r2_yellow_5(X0,X1,X1)) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v3_orders_2(X0)) [flattening 47]
49. ? [X0] : (? [X1] : ((k3_yellow_0(X0) = X1 | ~r2_yellow_5(X0,X1,X1)) & (k3_yellow_0(X0) != X1 | r2_yellow_5(X0,X1,X1)) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v3_orders_2(X0)) => (? [X1] : ((k3_yellow_0(sK0) = X1 | ~r2_yellow_5(sK0,X1,X1)) & (k3_yellow_0(sK0) != X1 | r2_yellow_5(sK0,X1,X1)) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v2_lattice3(sK0) & v1_lattice3(sK0) & v5_orders_2(sK0) & v3_orders_2(sK0)) [choice axiom]
50. ? [X1] : ((k3_yellow_0(sK0) = X1 | ~r2_yellow_5(sK0,X1,X1)) & (k3_yellow_0(sK0) != X1 | r2_yellow_5(sK0,X1,X1)) & m1_subset_1(X1,u1_struct_0(sK0))) => ((k3_yellow_0(sK0) = sK1 | ~r2_yellow_5(sK0,sK1,sK1)) & (k3_yellow_0(sK0) != sK1 | r2_yellow_5(sK0,sK1,sK1)) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
51. ((k3_yellow_0(sK0) = sK1 | ~r2_yellow_5(sK0,sK1,sK1)) & (k3_yellow_0(sK0) != sK1 | r2_yellow_5(sK0,sK1,sK1)) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v2_lattice3(sK0) & v1_lattice3(sK0) & v5_orders_2(sK0) & v3_orders_2(sK0) [skolemisation 48,50,49]
52. ! [X0] : (! [X1] : (! [X2] : (((r1_yellow_5(X0,X1,X2) | k3_yellow_0(X0) = k11_lattice3(X0,X1,X2)) & (k3_yellow_0(X0) != k11_lattice3(X0,X1,X2) | ~r1_yellow_5(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [nnf transformation 32]
55. ! [X0,X1,X2] : (((r2_yellow_5(X0,X1,X2) | ~r1_yellow_5(X0,X1,X2)) & (r1_yellow_5(X0,X1,X2) | ~r2_yellow_5(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0)) [nnf transformation 44]
58. v3_orders_2(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
59. v5_orders_2(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
61. v2_lattice3(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
62. l1_orders_2(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
63. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 51] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
64. k3_yellow_0(sK0) != sK1 | r2_yellow_5(sK0,sK1,sK1) [cnf transformation 51] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. k3_yellow_0(sK0) = sK1 | ~r2_yellow_5(sK0,sK1,sK1) [cnf transformation 51] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
68. ~v2_struct_0(X0) | ~v2_lattice3(X0) | ~l1_orders_2(X0) [cnf transformation 30] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
69. k3_yellow_0(X0) != k11_lattice3(X0,X1,X2) | ~r1_yellow_5(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0) [cnf transformation 52] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
70. k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) | r1_yellow_5(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0) [cnf transformation 52] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
71. k12_lattice3(X0,X1,X1) = X1 | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v3_orders_2(X0) [cnf transformation 34] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
74. k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 38] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
77. ~r2_yellow_5(X0,X1,X2) | r1_yellow_5(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 55] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
78. ~r1_yellow_5(X0,X1,X2) | r2_yellow_5(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 55] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
82. 1 <=> r2_yellow_5(sK0,sK1,sK1) [avatar definition]
83. r2_yellow_5(sK0,sK1,sK1) <- (1) [avatar component clause 82] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. 2 <=> k3_yellow_0(sK0) = sK1 [avatar definition]
87. k3_yellow_0(sK0) != sK1 <- (~2) [avatar component clause 86] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
88. k3_yellow_0(sK0) = sK1 <- (2) [avatar component clause 86] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. ~1 | 2 [avatar split clause 65,86,82]
90. 1 | ~2 [avatar split clause 64,86,82]
92. 3 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
94. m1_subset_1(sK1,u1_struct_0(sK0)) <- (3) [avatar component clause 92] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. 3 [avatar split clause 63,92]
97. 4 <=> l1_orders_2(sK0) [avatar definition]
99. l1_orders_2(sK0) <- (4) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. 4 [avatar split clause 62,97]
102. 5 <=> v2_lattice3(sK0) [avatar definition]
104. v2_lattice3(sK0) <- (5) [avatar component clause 102] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 5 [avatar split clause 61,102]
112. 7 <=> v5_orders_2(sK0) [avatar definition]
114. v5_orders_2(sK0) <- (7) [avatar component clause 112] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
115. 7 [avatar split clause 59,112]
117. 8 <=> v3_orders_2(sK0) [avatar definition]
119. v3_orders_2(sK0) <- (8) [avatar component clause 117] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
120. 8 [avatar split clause 58,117]
132. k3_yellow_0(X3) = k12_lattice3(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) | r1_yellow_5(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | v2_struct_0(X3) [superposition 74,70] {a:1,w:37,wCS:14800,sine:255,thAx:0,allAx:2,thDist:-2}
134. k3_yellow_0(X3) != k12_lattice3(X3,X4,X5) | ~r1_yellow_5(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | v2_struct_0(X3) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) [superposition 69,74] {a:1,w:37,wCS:14800,sine:255,thAx:0,allAx:2,thDist:-2}
137. k3_yellow_0(X3) != k12_lattice3(X3,X4,X5) | ~r1_yellow_5(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | v2_struct_0(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) [duplicate literal removal 134] {a:1,w:27,wCS:10800,sine:255,thAx:0,allAx:2,thDist:-2}
139. k3_yellow_0(X3) = k12_lattice3(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) | r1_yellow_5(X3,X4,X5) | v2_struct_0(X3) [duplicate literal removal 132] {a:1,w:27,wCS:10800,sine:255,thAx:0,allAx:2,thDist:-2}
140. k3_yellow_0(X3) = k12_lattice3(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) | r1_yellow_5(X3,X4,X5) [subsumption resolution 139,68] {a:1,w:25,wCS:10000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
142. k3_yellow_0(X3) != k12_lattice3(X3,X4,X5) | ~r1_yellow_5(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) [subsumption resolution 137,68] {a:1,w:25,wCS:10000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
150. k3_yellow_0(X0) = X1 | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v3_orders_2(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | r1_yellow_5(X0,X1,X1) [superposition 71,140] {a:2,w:34,wCS:13600,sine:1,thAx:0,allAx:3,thDist:-3}
157. ~m1_subset_1(X1,u1_struct_0(X0)) | k3_yellow_0(X0) = X1 | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v3_orders_2(X0) | r1_yellow_5(X0,X1,X1) [duplicate literal removal 150] {a:2,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
161. k3_yellow_0(X0) != X1 | ~r1_yellow_5(X0,X1,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v3_orders_2(X0) [superposition 142,71] {a:2,w:34,wCS:13600,sine:1,thAx:0,allAx:3,thDist:-3}
168. k3_yellow_0(X0) != X1 | ~r1_yellow_5(X0,X1,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v3_orders_2(X0) [duplicate literal removal 161] {a:2,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
172. k3_yellow_0(sK0) = sK1 | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) | ~v3_orders_2(sK0) | r1_yellow_5(sK0,sK1,sK1) <- (3) [resolution 157,94] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
177. ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) | ~v3_orders_2(sK0) | r1_yellow_5(sK0,sK1,sK1) <- (~2, 3) [subsumption resolution 172,87] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
178. ~v2_lattice3(sK0) | ~v5_orders_2(sK0) | ~v3_orders_2(sK0) | r1_yellow_5(sK0,sK1,sK1) <- (~2, 3, 4) [subsumption resolution 177,99] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
179. ~v5_orders_2(sK0) | ~v3_orders_2(sK0) | r1_yellow_5(sK0,sK1,sK1) <- (~2, 3, 4, 5) [subsumption resolution 178,104] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
180. ~v3_orders_2(sK0) | r1_yellow_5(sK0,sK1,sK1) <- (~2, 3, 4, 5, 7) [subsumption resolution 179,114] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
181. r1_yellow_5(sK0,sK1,sK1) <- (~2, 3, 4, 5, 7, 8) [subsumption resolution 180,119] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
183. 10 <=> r1_yellow_5(sK0,sK1,sK1) [avatar definition]
185. r1_yellow_5(sK0,sK1,sK1) <- (10) [avatar component clause 183] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
186. 10 | 2 | ~3 | ~4 | ~5 | ~7 | ~8 [avatar split clause 181,117,112,102,97,92,86,183]
188. r2_yellow_5(sK0,sK1,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (10) [resolution 185,78] {a:4,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
189. r2_yellow_5(sK0,sK1,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (10) [duplicate literal removal 188] {a:4,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
197. r2_yellow_5(sK0,sK1,sK1) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (3, 10) [subsumption resolution 189,94] {a:4,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
198. r2_yellow_5(sK0,sK1,sK1) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (3, 4, 10) [subsumption resolution 197,99] {a:4,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
199. r2_yellow_5(sK0,sK1,sK1) | ~v5_orders_2(sK0) <- (3, 4, 5, 10) [subsumption resolution 198,104] {a:4,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
200. r2_yellow_5(sK0,sK1,sK1) <- (3, 4, 5, 7, 10) [subsumption resolution 199,114] {a:4,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
201. 1 | ~3 | ~4 | ~5 | ~7 | ~10 [avatar split clause 200,183,112,102,97,92,82]
202. r1_yellow_5(sK0,sK1,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (1) [resolution 83,77] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
205. r1_yellow_5(sK0,sK1,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (1) [duplicate literal removal 202] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
207. $false <- (2, 3, 4, 5, 7, 8, 10) [unit resulting resolution 119,114,104,99,94,185,88,168] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:27,thDist:-27}
210. ~2 | ~3 | ~4 | ~5 | ~7 | ~8 | ~10 [avatar contradiction clause 207]
216. r1_yellow_5(sK0,sK1,sK1) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (1, 3) [subsumption resolution 205,94] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
217. r1_yellow_5(sK0,sK1,sK1) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (1, 3, 4) [subsumption resolution 216,99] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
218. r1_yellow_5(sK0,sK1,sK1) | ~v5_orders_2(sK0) <- (1, 3, 4, 5) [subsumption resolution 217,104] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
219. r1_yellow_5(sK0,sK1,sK1) <- (1, 3, 4, 5, 7) [subsumption resolution 218,114] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
220. 10 | ~1 | ~3 | ~4 | ~5 | ~7 [avatar split clause 219,112,102,97,92,82,183]
221. $false [avatar sat refutation 89,90,95,100,105,115,120,186,201,210,220] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:132,thDist:-132}
% SZS output end Proof for yellow_5__t15_yellow_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 19
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 127
% Active clauses: 30
% Passive clauses: 44
% Final active clauses: 27
% Final passive clauses: 12
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Duplicate literals: 97
% Fw subsumption resolutions: 26
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 8
% Unit resulting resolution: 2
% Forward superposition: 10
% Backward superposition: 12
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 7
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4989
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% non unit clause index maintenance: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

