i: [1,48,0,3,1,1,0,0] t121_zfmisc_1
a: [3,78,0,3,1,48]
s: 78
i: [1,47,0,3,1,1,0,0] t121_zfmisc_1
a: [3,83,0,3,1,47]
s: 83
i: [1,50,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,88,0,2,1,50]
s: 88
i: [1,51,0,2,1,0,0,2] fc1_xboole_0
a: [3,93,0,2,1,51]
s: 93
i: [1,68,0,2,1,0,0,2] rc2_xboole_0
a: [3,103,0,2,1,68]
s: 103
i: [1,69,0,2,1,0,0,2] rc1_xboole_0
a: [3,108,0,2,1,69]
s: 108
i: [1,49,0,11,1,1,0,0] t121_zfmisc_1
a: [3,73,0,11,1,49]
s: 73
i: [1,52,0,3,1,0,0,255] d2_xboole_0
a: [3,98,0,3,1,52]
s: 98
r: 88
i: [1,61,0,4,1,0,0,1] fc3_xboole_0
s: 61
i: [1,53,0,5,1,0,0,1] t4_boole
s: 53
i: [1,54,0,5,1,0,0,1] t3_boole
s: 54
i: [1,55,0,5,1,0,0,2] t1_boole
s: 55
i: [1,56,0,5,2,0,0,2] t6_boole
s: 56
i: [1,59,0,5,2,0,0,1] t7_boole
s: 59
i: [1,62,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 62
i: [1,65,0,6,2,0,0,1] fc4_xboole_0
s: 65
i: [1,66,0,6,2,0,0,1] fc5_xboole_0
s: 66
i: [1,67,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 67
i: [1,58,0,7,3,0,0,2] t8_boole
s: 58
i: [1,63,0,7,1,0,0,1] commutativity_k2_tarski
s: 63
i: [1,64,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 64
d: [2,122,1,5,1,0,44,64,55]
s: 122
i: [1,57,0,9,1,0,0,1] t40_xboole_1
s: 57
d: [2,141,1,7,1,1,44,73,57]
a: [3,148,1,7,1,141]
s: 148
i: [1,60,0,13,3,0,0,1] t120_zfmisc_1
s: 60
d: [2,151,2,0,0,3,52,83,78,148,60]
e: 151
d: [2,157,2,9,3,1,44,148,60]
d: [2,158,2,6,2,1,19,157]
d: [2,160,2,3,1,2,21,158,78]
d: [2,161,2,0,0,3,21,160,83]
e: 161
f: 161,141,47,48,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t121_zfmisc_1
% SZS output start Proof for zfmisc_1__t121_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ~(k4_xboole_0(k2_xboole_0(X2,k2_tarski(X0,X1)),k2_tarski(X0,X1)) != X2 & ~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) [input t121_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ~(k4_xboole_0(k2_xboole_0(X2,k2_tarski(X0,X1)),k2_tarski(X0,X1)) != X2 & ~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) [negated conjecture 1]
19. ! [X0] : ! [X1] : ! [X2] : ~(k4_xboole_0(X2,k2_tarski(X0,X1)) != X2 & ~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) [input t120_zfmisc_1]
22. ! [X0] : ! [X1] : k4_xboole_0(k2_xboole_0(X0,X1),X1) = k4_xboole_0(X0,X1) [input t40_xboole_1]
27. ~! [X0,X1,X2] : ~(k4_xboole_0(k2_xboole_0(X2,k2_tarski(X0,X1)),k2_tarski(X0,X1)) != X2 & ~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) [flattening 2]
28. ! [X0,X1] : k4_xboole_0(k2_xboole_0(X0,X1),X1) = k4_xboole_0(X0,X1) [flattening 22]
31. ! [X0,X1,X2] : ~(k4_xboole_0(X2,k2_tarski(X0,X1)) != X2 & ~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) [flattening 19]
33. ? [X0,X1,X2] : (k4_xboole_0(k2_xboole_0(X2,k2_tarski(X0,X1)),k2_tarski(X0,X1)) != X2 & ~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) [ennf transformation 27]
37. ! [X0,X1,X2] : (k4_xboole_0(X2,k2_tarski(X0,X1)) = X2 | r2_hidden(X1,X2) | r2_hidden(X0,X2)) [ennf transformation 31]
41. ? [X0,X1,X2] : (k4_xboole_0(k2_xboole_0(X2,k2_tarski(X0,X1)),k2_tarski(X0,X1)) != X2 & ~r2_hidden(X1,X2) & ~r2_hidden(X0,X2)) => (sK2 != k4_xboole_0(k2_xboole_0(sK2,k2_tarski(sK0,sK1)),k2_tarski(sK0,sK1)) & ~r2_hidden(sK1,sK2) & ~r2_hidden(sK0,sK2)) [choice axiom]
42. sK2 != k4_xboole_0(k2_xboole_0(sK2,k2_tarski(sK0,sK1)),k2_tarski(sK0,sK1)) & ~r2_hidden(sK1,sK2) & ~r2_hidden(sK0,sK2) [skolemisation 33,41]
47. ~r2_hidden(sK0,sK2) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. ~r2_hidden(sK1,sK2) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. sK2 != k4_xboole_0(k2_xboole_0(sK2,k2_tarski(sK0,sK1)),k2_tarski(sK0,sK1)) [cnf transformation 42] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
57. k4_xboole_0(k2_xboole_0(X0,X1),X1) = k4_xboole_0(X0,X1) [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. k4_xboole_0(X2,k2_tarski(X0,X1)) = X2 | r2_hidden(X1,X2) | r2_hidden(X0,X2) [cnf transformation 37] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. 1 <=> sK2 = k4_xboole_0(k2_xboole_0(sK2,k2_tarski(sK0,sK1)),k2_tarski(sK0,sK1)) [avatar definition]
73. sK2 != k4_xboole_0(k2_xboole_0(sK2,k2_tarski(sK0,sK1)),k2_tarski(sK0,sK1)) <- (~1) [avatar component clause 71] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
74. ~1 [avatar split clause 49,71]
76. 2 <=> r2_hidden(sK1,sK2) [avatar definition]
78. ~r2_hidden(sK1,sK2) <- (~2) [avatar component clause 76] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. ~2 [avatar split clause 48,76]
81. 3 <=> r2_hidden(sK0,sK2) [avatar definition]
83. ~r2_hidden(sK0,sK2) <- (~3) [avatar component clause 81] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. ~3 [avatar split clause 47,81]
141. sK2 != k4_xboole_0(sK2,k2_tarski(sK0,sK1)) <- (~1) [superposition 73,57] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
146. 9 <=> sK2 = k4_xboole_0(sK2,k2_tarski(sK0,sK1)) [avatar definition]
148. sK2 != k4_xboole_0(sK2,k2_tarski(sK0,sK1)) <- (~9) [avatar component clause 146] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
149. ~9 | 1 [avatar split clause 141,71,146]
157. sK2 != sK2 | r2_hidden(sK1,sK2) | r2_hidden(sK0,sK2) <- (~9) [superposition 148,60] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
158. r2_hidden(sK1,sK2) | r2_hidden(sK0,sK2) <- (~9) [trivial inequality removal 157] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
160. r2_hidden(sK0,sK2) <- (~2, ~9) [subsumption resolution 158,78] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
161. $false <- (~2, ~3, ~9) [subsumption resolution 160,83] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
162. 2 | 3 | 9 [avatar contradiction clause 161]
163. $false [avatar sat refutation 74,79,84,149,162] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:38,thDist:-38}
% SZS output end Proof for zfmisc_1__t121_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 78
% Active clauses: 25
% Passive clauses: 36
% Final active clauses: 24
% Final passive clauses: 11
% Main loop iterations started: 25
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 2
% Fw demodulations: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 17
% 
% >>> Generating Inferences
% Unit resulting resolution: 1
% Forward superposition: 16
% Backward superposition: 21
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 8
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

