i: [1,101,0,2,1,1,0,0] t3_pzfmisc1
a: [3,144,0,2,1,101]
s: 144
i: [1,100,0,2,1,1,0,0] t3_pzfmisc1
a: [3,149,0,2,1,100]
s: 149
i: [1,98,0,2,1,1,0,0] t3_pzfmisc1
a: [3,159,0,2,1,98]
s: 159
i: [1,96,0,2,1,1,0,0] t3_pzfmisc1
a: [3,169,0,2,1,96]
s: 169
i: [1,95,0,2,1,1,0,0] t3_pzfmisc1
a: [3,174,0,2,1,95]
s: 174
i: [1,99,0,3,1,1,0,0] t3_pzfmisc1
a: [3,154,0,3,1,99]
s: 154
i: [1,97,0,3,1,1,0,0] t3_pzfmisc1
a: [3,164,0,3,1,97]
s: 164
i: [1,102,0,2,1,0,0,255] fc1_xboole_0
a: [3,179,0,2,1,102]
s: 179
i: [1,128,0,2,1,0,0,1] rc2_xboole_0
a: [3,184,0,2,1,128]
s: 184
i: [1,129,0,2,1,0,0,1] rc1_xboole_0
a: [3,189,0,2,1,129]
s: 189
i: [1,130,0,2,1,0,0,255] rc1_finset_1
a: [3,194,0,2,1,130]
s: 194
i: [1,131,0,2,1,0,0,255] rc8_finset_1
a: [3,199,0,2,1,131]
s: 199
i: [1,134,0,2,1,0,0,255] rc9_finset_1
a: [3,204,0,2,1,134]
s: 204
i: [1,133,0,2,1,0,0,255] rc9_finset_1
a: [3,209,0,2,1,133]
s: 209
i: [1,132,0,2,1,0,0,255] rc9_finset_1
a: [3,214,0,2,1,132]
s: 214
i: [1,137,0,2,1,0,0,255] rc3_finset_1
a: [3,219,0,2,1,137]
s: 219
i: [1,136,0,2,1,0,0,255] rc3_finset_1
a: [3,224,0,2,1,136]
s: 224
i: [1,135,0,2,1,0,0,255] rc3_finset_1
a: [3,229,0,2,1,135]
s: 229
i: [1,140,0,2,1,0,0,3] rc1_funcop_1
a: [3,234,0,2,1,140]
s: 234
i: [1,139,0,2,1,0,0,3] rc1_funcop_1
a: [3,239,0,2,1,139]
s: 239
i: [1,138,0,2,1,0,0,3] rc1_funcop_1
a: [3,244,0,2,1,138]
s: 244
i: [1,120,0,3,1,0,0,1] rc4_funcop_1
s: 120
i: [1,121,0,3,1,0,0,1] rc4_funcop_1
s: 121
i: [1,123,0,3,1,0,0,1] rc4_funcop_1
s: 123
i: [1,122,0,4,1,0,0,1] rc4_funcop_1
s: 122
i: [1,124,0,4,1,0,0,1] rc4_funcop_1
s: 124
i: [1,103,0,5,2,0,0,255] t6_boole
s: 103
i: [1,110,0,5,2,0,0,2] d1_xboole_0
s: 110
i: [1,119,0,5,2,0,0,2] t7_boole
s: 119
i: [1,111,0,6,2,0,0,2] d1_xboole_0
s: 111
i: [1,125,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 125
d: [2,249,1,6,2,0,40,125,111]
s: 249
i: [1,118,0,7,3,0,0,255] t8_boole
s: 118
d: [2,246,1,7,3,0,44,103,103]
s: 246
i: [1,109,0,8,4,0,0,3] cc3_funcop_1
s: 109
i: [1,126,0,10,4,0,0,2] fc8_funcop_1
s: 126
d: [2,250,1,6,3,1,40,109,159]
d: [2,255,1,4,2,2,21,250,169]
a: [3,263,1,2,1,255]
s: 263
i: [1,127,0,10,4,0,0,2] fc8_funcop_1
s: 127
d: [2,251,1,9,3,0,40,109,123]
d: [2,265,1,6,2,0,21,251,120]
s: 265
i: [1,113,0,17,6,0,0,1] d12_pboole
s: 113
i: [1,116,0,17,6,0,0,1] d13_pboole
s: 116
d: [2,266,1,18,6,0,40,127,109]
d: [2,267,1,14,5,0,21,266,126]
s: 267
i: [1,114,0,18,6,0,0,1] d12_pboole
s: 114
d: [2,269,1,14,6,0,40,113,110]
s: 269
i: [1,117,0,18,6,0,0,1] d13_pboole
s: 117
i: [1,112,0,19,7,0,0,1] d12_pboole
s: 112
d: [2,273,2,15,5,0,40,269,124]
d: [2,274,2,12,4,0,21,273,123]
d: [2,275,2,8,3,0,21,274,122]
d: [2,276,2,5,2,0,21,275,120]
s: 276
i: [1,115,0,19,7,0,0,1] d13_pboole
s: 115
d: [2,271,1,14,6,0,40,116,110]
s: 271
d: [2,268,1,17,6,0,40,113,125]
s: 268
d: [2,270,1,17,6,0,40,116,125]
s: 270
d: [2,277,2,24,9,0,40,117,267]
d: [2,278,2,20,7,0,18,277]
s: 278
d: [2,284,1,30,12,0,40,115,112]
d: [2,285,1,26,10,0,18,284]
s: 285
d: [2,279,2,25,10,0,40,112,267]
d: [2,282,2,21,8,0,18,279]
s: 282
d: [2,289,2,26,10,0,40,285,111]
s: 289
d: [2,290,2,37,14,0,40,285,113]
s: 290
d: [2,292,3,24,10,0,40,289,111]
d: [2,293,3,16,7,0,18,292]
s: 293
r: 289
d: [2,296,4,0,0,7,52,174,149,144,169,159,164,154,293]
e: 296
d: [2,297,4,13,6,1,40,293,154]
d: [2,300,4,11,5,2,21,297,159]
d: [2,301,4,8,4,3,21,300,164]
d: [2,302,4,6,3,4,21,301,169]
d: [2,303,4,4,2,5,21,302,144]
d: [2,304,4,2,1,6,21,303,149]
d: [2,305,4,0,0,7,21,304,174]
e: 305
f: 305,95,96,97,98,99,100,101
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pzfmisc1__t3_pzfmisc1
% SZS output start Proof for pzfmisc1__t3_pzfmisc1
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ~(v2_relat_1(X1) & v3_relat_1(X1)))) [input t3_pzfmisc1]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ~(v2_relat_1(X1) & v3_relat_1(X1)))) [negated conjecture 1]
11. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v3_relat_1(X1) <=> ! [X2] : (r2_hidden(X2,X0) => v1_xboole_0(k1_funct_1(X1,X2))))) [input d12_pboole]
12. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v2_relat_1(X1) <=> ! [X2] : ~(v1_xboole_0(k1_funct_1(X1,X2)) & r2_hidden(X2,X0)))) [input d13_pboole]
13. ! [X0] : (v1_xboole_0(X0) <=> ! [X1] : ~r2_hidden(X1,X0)) [input d1_xboole_0]
30. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v3_relat_1(X1) <=> ! [X2] : (r2_hidden(X2,X0) => v1_xboole_0(k1_funct_1(X1,X2))))) [flattening 11]
31. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v2_relat_1(X1) <=> ! [X2] : ~(v1_xboole_0(k1_funct_1(X1,X2)) & r2_hidden(X2,X0)))) [flattening 12]
46. ? [X0] : (? [X1] : ((v2_relat_1(X1) & v3_relat_1(X1)) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) & ~v1_xboole_0(X0)) [ennf transformation 2]
47. ? [X0] : (? [X1] : (v2_relat_1(X1) & v3_relat_1(X1) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) & ~v1_xboole_0(X0)) [flattening 46]
55. ! [X0,X1] : ((v3_relat_1(X1) <=> ! [X2] : (v1_xboole_0(k1_funct_1(X1,X2)) | ~r2_hidden(X2,X0))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 30]
56. ! [X0,X1] : ((v3_relat_1(X1) <=> ! [X2] : (v1_xboole_0(k1_funct_1(X1,X2)) | ~r2_hidden(X2,X0))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 55]
57. ! [X0,X1] : ((v2_relat_1(X1) <=> ! [X2] : (~v1_xboole_0(k1_funct_1(X1,X2)) | ~r2_hidden(X2,X0))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 31]
58. ! [X0,X1] : ((v2_relat_1(X1) <=> ! [X2] : (~v1_xboole_0(k1_funct_1(X1,X2)) | ~r2_hidden(X2,X0))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 57]
64. ? [X0] : (? [X1] : (v2_relat_1(X1) & v3_relat_1(X1) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) & ~v1_xboole_0(X0)) => (? [X1] : (v2_relat_1(X1) & v3_relat_1(X1) & v1_partfun1(X1,sK0) & v1_funct_1(X1) & v4_relat_1(X1,sK0) & v1_relat_1(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
65. ? [X1] : (v2_relat_1(X1) & v3_relat_1(X1) & v1_partfun1(X1,sK0) & v1_funct_1(X1) & v4_relat_1(X1,sK0) & v1_relat_1(X1)) => (v2_relat_1(sK1) & v3_relat_1(sK1) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
66. (v2_relat_1(sK1) & v3_relat_1(sK1) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) & ~v1_xboole_0(sK0) [skolemisation 47,65,64]
67. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X1] : ~r2_hidden(X1,X0) | ~v1_xboole_0(X0))) [nnf transformation 13]
68. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [rectify 67]
69. ! [X0] : (? [X1] : r2_hidden(X1,X0) => r2_hidden(sK2(X0),X0)) [choice axiom]
70. ! [X0] : ((v1_xboole_0(X0) | r2_hidden(sK2(X0),X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [skolemisation 68,69]
71. ! [X0,X1] : (((v3_relat_1(X1) | ? [X2] : (~v1_xboole_0(k1_funct_1(X1,X2)) & r2_hidden(X2,X0))) & (! [X2] : (v1_xboole_0(k1_funct_1(X1,X2)) | ~r2_hidden(X2,X0)) | ~v3_relat_1(X1))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 56]
72. ! [X0,X1] : (((v3_relat_1(X1) | ? [X2] : (~v1_xboole_0(k1_funct_1(X1,X2)) & r2_hidden(X2,X0))) & (! [X3] : (v1_xboole_0(k1_funct_1(X1,X3)) | ~r2_hidden(X3,X0)) | ~v3_relat_1(X1))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [rectify 71]
73. ! [X1,X0] : (? [X2] : (~v1_xboole_0(k1_funct_1(X1,X2)) & r2_hidden(X2,X0)) => (~v1_xboole_0(k1_funct_1(X1,sK3(X0,X1))) & r2_hidden(sK3(X0,X1),X0))) [choice axiom]
74. ! [X0,X1] : (((v3_relat_1(X1) | (~v1_xboole_0(k1_funct_1(X1,sK3(X0,X1))) & r2_hidden(sK3(X0,X1),X0))) & (! [X3] : (v1_xboole_0(k1_funct_1(X1,X3)) | ~r2_hidden(X3,X0)) | ~v3_relat_1(X1))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [skolemisation 72,73]
75. ! [X0,X1] : (((v2_relat_1(X1) | ? [X2] : (v1_xboole_0(k1_funct_1(X1,X2)) & r2_hidden(X2,X0))) & (! [X2] : (~v1_xboole_0(k1_funct_1(X1,X2)) | ~r2_hidden(X2,X0)) | ~v2_relat_1(X1))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 58]
76. ! [X0,X1] : (((v2_relat_1(X1) | ? [X2] : (v1_xboole_0(k1_funct_1(X1,X2)) & r2_hidden(X2,X0))) & (! [X3] : (~v1_xboole_0(k1_funct_1(X1,X3)) | ~r2_hidden(X3,X0)) | ~v2_relat_1(X1))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [rectify 75]
77. ! [X1,X0] : (? [X2] : (v1_xboole_0(k1_funct_1(X1,X2)) & r2_hidden(X2,X0)) => (v1_xboole_0(k1_funct_1(X1,sK4(X0,X1))) & r2_hidden(sK4(X0,X1),X0))) [choice axiom]
78. ! [X0,X1] : (((v2_relat_1(X1) | (v1_xboole_0(k1_funct_1(X1,sK4(X0,X1))) & r2_hidden(sK4(X0,X1),X0))) & (! [X3] : (~v1_xboole_0(k1_funct_1(X1,X3)) | ~r2_hidden(X3,X0)) | ~v2_relat_1(X1))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [skolemisation 76,77]
95. ~v1_xboole_0(sK0) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
96. v1_relat_1(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
97. v4_relat_1(sK1,sK0) [cnf transformation 66] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
98. v1_funct_1(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
99. v1_partfun1(sK1,sK0) [cnf transformation 66] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. v3_relat_1(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
101. v2_relat_1(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
111. r2_hidden(sK2(X0),X0) | v1_xboole_0(X0) [cnf transformation 70] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
112. v1_xboole_0(k1_funct_1(X1,X3)) | ~r2_hidden(X3,X0) | ~v3_relat_1(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 74] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
115. ~v1_xboole_0(k1_funct_1(X1,X3)) | ~r2_hidden(X3,X0) | ~v2_relat_1(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 78] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
142. 1 <=> v2_relat_1(sK1) [avatar definition]
144. v2_relat_1(sK1) <- (1) [avatar component clause 142] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
145. 1 [avatar split clause 101,142]
147. 2 <=> v3_relat_1(sK1) [avatar definition]
149. v3_relat_1(sK1) <- (2) [avatar component clause 147] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
150. 2 [avatar split clause 100,147]
152. 3 <=> v1_partfun1(sK1,sK0) [avatar definition]
154. v1_partfun1(sK1,sK0) <- (3) [avatar component clause 152] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
155. 3 [avatar split clause 99,152]
157. 4 <=> v1_funct_1(sK1) [avatar definition]
159. v1_funct_1(sK1) <- (4) [avatar component clause 157] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
160. 4 [avatar split clause 98,157]
162. 5 <=> v4_relat_1(sK1,sK0) [avatar definition]
164. v4_relat_1(sK1,sK0) <- (5) [avatar component clause 162] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
165. 5 [avatar split clause 97,162]
167. 6 <=> v1_relat_1(sK1) [avatar definition]
169. v1_relat_1(sK1) <- (6) [avatar component clause 167] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
170. 6 [avatar split clause 96,167]
172. 7 <=> v1_xboole_0(sK0) [avatar definition]
174. ~v1_xboole_0(sK0) <- (~7) [avatar component clause 172] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
175. ~7 [avatar split clause 95,172]
284. ~r2_hidden(X3,X4) | ~v2_relat_1(X5) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~r2_hidden(X3,X6) | ~v3_relat_1(X5) | ~v1_partfun1(X5,X6) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X6) | ~v1_relat_1(X5) [resolution 115,112] {a:1,w:30,wCS:12000,sine:1,thAx:0,allAx:4,thDist:-4}
285. ~r2_hidden(X3,X6) | ~v2_relat_1(X5) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~r2_hidden(X3,X4) | ~v3_relat_1(X5) | ~v1_partfun1(X5,X6) | ~v4_relat_1(X5,X6) [duplicate literal removal 284] {a:1,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
289. ~r2_hidden(sK2(X2),X1) | ~v1_partfun1(X0,X1) | ~v1_funct_1(X0) | ~v4_relat_1(X0,X1) | ~v1_relat_1(X0) | ~v2_relat_1(X0) | ~v3_relat_1(X0) | ~v1_partfun1(X0,X2) | ~v4_relat_1(X0,X2) | v1_xboole_0(X2) [resolution 285,111] {a:2,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
292. ~v1_partfun1(X0,X1) | ~v1_funct_1(X0) | ~v4_relat_1(X0,X1) | ~v1_relat_1(X0) | ~v2_relat_1(X0) | ~v3_relat_1(X0) | ~v1_partfun1(X0,X1) | ~v4_relat_1(X0,X1) | v1_xboole_0(X1) | v1_xboole_0(X1) [resolution 289,111] {a:3,w:24,wCS:9600,sine:1,thAx:0,allAx:8,thDist:-8}
293. ~v1_partfun1(X0,X1) | ~v1_funct_1(X0) | ~v4_relat_1(X0,X1) | ~v1_relat_1(X0) | ~v2_relat_1(X0) | ~v3_relat_1(X0) | v1_xboole_0(X1) [duplicate literal removal 292] {a:3,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:8,thDist:-8}
297. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v2_relat_1(sK1) | ~v3_relat_1(sK1) | v1_xboole_0(sK0) <- (3) [resolution 293,154] {a:4,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
300. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v2_relat_1(sK1) | ~v3_relat_1(sK1) | v1_xboole_0(sK0) <- (3, 4) [subsumption resolution 297,159] {a:4,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
301. ~v1_relat_1(sK1) | ~v2_relat_1(sK1) | ~v3_relat_1(sK1) | v1_xboole_0(sK0) <- (3, 4, 5) [subsumption resolution 300,164] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
302. ~v2_relat_1(sK1) | ~v3_relat_1(sK1) | v1_xboole_0(sK0) <- (3, 4, 5, 6) [subsumption resolution 301,169] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
303. ~v3_relat_1(sK1) | v1_xboole_0(sK0) <- (1, 3, 4, 5, 6) [subsumption resolution 302,144] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
304. v1_xboole_0(sK0) <- (1, 2, 3, 4, 5, 6) [subsumption resolution 303,149] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
305. $false <- (1, 2, 3, 4, 5, 6, ~7) [subsumption resolution 304,174] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
306. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | 7 [avatar contradiction clause 305]
311. $false [avatar sat refutation 145,150,155,160,165,170,175,306] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:117,thDist:-117}
% SZS output end Proof for pzfmisc1__t3_pzfmisc1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 28
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 5
% 
% >>> Saturation
% Initial clauses: 46
% Generated clauses: 122
% Active clauses: 55
% Passive clauses: 59
% Final active clauses: 54
% Final passive clauses: 2
% Main loop iterations started: 57
% 
% >>> Simplifying Inferences
% Duplicate literals: 20
% Fw subsumption resolutions: 16
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 8
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 28
% Unit resulting resolution: 1
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 21
% 
% Memory used [KB]: 9466
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% property evaluation: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% resolution: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

