i: [1,29,0,11,1,1,0,0] t94_xboole_1
a: [3,44,0,11,1,29]
s: 44
i: [1,35,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 35
i: [1,36,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 36
i: [1,40,0,6,2,0,0,255] symmetry_r1_xboole_0
s: 40
i: [1,30,0,7,1,0,0,255] l97_xboole_1
s: 30
i: [1,37,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 37
d: [2,47,1,5,1,0,44,30,35]
s: 47
i: [1,38,0,7,1,0,0,1] commutativity_k5_xboole_0
s: 38
d: [2,50,2,5,1,0,40,47,40]
s: 50
i: [1,39,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 39
i: [1,33,0,8,2,0,0,255] t83_xboole_1
s: 33
d: [2,46,1,7,1,0,40,30,40]
s: 46
i: [1,34,0,8,2,0,0,255] t83_xboole_1
s: 34
d: [2,48,1,7,1,0,44,30,37]
s: 48
i: [1,31,0,11,1,0,0,255] d6_xboole_0
s: 31
i: [1,32,0,11,1,0,0,1] t93_xboole_1
s: 32
d: [2,49,1,7,1,0,44,30,37]
s: 49
d: [2,51,1,7,1,0,44,30,38]
s: 51
d: [2,52,1,7,1,0,44,30,38]
s: 52
d: [2,69,1,7,1,0,44,31,36]
s: 69
d: [2,72,1,7,1,0,44,36,31]
s: 72
d: [2,78,1,9,1,0,44,32,35]
d: [2,81,1,7,1,0,22,78,36]
s: 81
d: [2,70,1,11,1,0,44,31,39]
s: 70
d: [2,55,2,7,1,0,44,46,38]
s: 55
d: [2,71,1,11,1,0,44,31,39]
s: 71
d: [2,73,1,11,1,0,44,39,31]
s: 73
d: [2,56,2,7,1,0,44,46,38]
s: 56
d: [2,74,1,11,1,0,44,39,31]
s: 74
d: [2,58,2,7,1,0,44,46,37]
s: 58
d: [2,76,1,11,1,0,44,32,38]
s: 76
d: [2,77,1,11,1,0,44,32,38]
s: 77
d: [2,59,2,7,1,0,44,46,37]
s: 59
d: [2,79,1,11,1,0,44,32,37]
s: 79
d: [2,61,2,7,1,0,40,48,40]
s: 61
d: [2,80,1,11,1,0,44,32,37]
s: 80
d: [2,67,1,12,2,0,44,31,33]
s: 67
d: [2,87,2,7,1,0,44,81,39]
s: 87
d: [2,68,1,12,2,0,44,31,33]
d: [2,75,1,12,2,0,22,68,39]
s: 75
d: [2,88,2,7,1,0,44,81,39]
s: 88
d: [2,89,2,7,1,0,44,39,81]
s: 89
d: [2,90,2,7,1,0,44,39,81]
s: 90
d: [2,82,2,8,2,0,44,69,33]
s: 82
d: [2,85,2,8,2,0,44,34,69]
s: 85
d: [2,86,2,8,2,0,44,33,69]
s: 86
d: [2,128,2,12,2,0,44,75,69]
d: [2,134,2,8,2,0,22,128,87]
s: 134
d: [2,124,2,13,3,0,44,67,33]
d: [2,126,2,10,2,0,21,124,40]
s: 126
d: [2,155,3,18,2,1,44,44,126]
d: [2,182,3,7,1,1,21,155,32]
d: [2,183,3,0,0,1,21,182,30]
e: 183
f: 183,29
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t94_xboole_1
% SZS output start Proof for xboole_1__t94_xboole_1
1. ! [X0] : ! [X1] : k2_xboole_0(X0,X1) = k5_xboole_0(k5_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [input t94_xboole_1]
2. ~! [X0] : ! [X1] : k2_xboole_0(X0,X1) = k5_xboole_0(k5_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [negated conjecture 1]
6. ! [X0] : ! [X1] : k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [input d6_xboole_0]
13. ! [X0] : ! [X1] : r1_xboole_0(k3_xboole_0(X0,X1),k5_xboole_0(X0,X1)) [input l97_xboole_1]
14. ! [X0,X1] : (r1_xboole_0(X0,X1) => r1_xboole_0(X1,X0)) [input symmetry_r1_xboole_0]
15. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [input t83_xboole_1]
16. ! [X0] : ! [X1] : k2_xboole_0(X0,X1) = k2_xboole_0(k5_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [input t93_xboole_1]
17. ~! [X0,X1] : k2_xboole_0(X0,X1) = k5_xboole_0(k5_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [flattening 2]
18. ! [X0,X1] : r1_xboole_0(k3_xboole_0(X0,X1),k5_xboole_0(X0,X1)) [flattening 13]
19. ! [X0,X1] : k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [flattening 6]
20. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(k5_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [flattening 16]
21. ! [X0,X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [flattening 15]
24. ? [X0,X1] : k2_xboole_0(X0,X1) != k5_xboole_0(k5_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [ennf transformation 17]
25. ! [X0,X1] : (r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 14]
26. ? [X0,X1] : k2_xboole_0(X0,X1) != k5_xboole_0(k5_xboole_0(X0,X1),k3_xboole_0(X0,X1)) => k2_xboole_0(sK0,sK1) != k5_xboole_0(k5_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) [choice axiom]
27. k2_xboole_0(sK0,sK1) != k5_xboole_0(k5_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) [skolemisation 24,26]
28. ! [X0,X1] : ((r1_xboole_0(X0,X1) | k4_xboole_0(X0,X1) != X0) & (k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1))) [nnf transformation 21]
29. k2_xboole_0(sK0,sK1) != k5_xboole_0(k5_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) [cnf transformation 27] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. r1_xboole_0(k3_xboole_0(X0,X1),k5_xboole_0(X0,X1)) [cnf transformation 18] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
31. k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
32. k2_xboole_0(X0,X1) = k2_xboole_0(k5_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [cnf transformation 20] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1) [cnf transformation 28] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. ~r1_xboole_0(X0,X1) | r1_xboole_0(X1,X0) [cnf transformation 25] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. 1 <=> k2_xboole_0(sK0,sK1) = k5_xboole_0(k5_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) [avatar definition]
44. k2_xboole_0(sK0,sK1) != k5_xboole_0(k5_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) <- (~1) [avatar component clause 42] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
45. ~1 [avatar split clause 29,42]
67. k5_xboole_0(X0,X1) = k2_xboole_0(X0,k4_xboole_0(X1,X0)) | ~r1_xboole_0(X0,X1) [superposition 31,33] {a:1,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
124. k2_xboole_0(X1,X0) = k5_xboole_0(X1,X0) | ~r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1) [superposition 67,33] {a:2,w:13,wCS:5200,sine:255,thAx:0,allAx:3,thDist:-3}
126. k2_xboole_0(X1,X0) = k5_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1) [subsumption resolution 124,40] {a:2,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
155. k2_xboole_0(sK0,sK1) != k2_xboole_0(k5_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) | ~r1_xboole_0(k3_xboole_0(sK0,sK1),k5_xboole_0(sK0,sK1)) <- (~1) [superposition 44,126] {a:3,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
182. ~r1_xboole_0(k3_xboole_0(sK0,sK1),k5_xboole_0(sK0,sK1)) <- (~1) [subsumption resolution 155,32] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
183. $false <- (~1) [subsumption resolution 182,30] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
184. 1 [avatar contradiction clause 183]
190. $false [avatar sat refutation 45,184] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
% SZS output end Proof for xboole_1__t94_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 147
% Active clauses: 28
% Passive clauses: 56
% Final active clauses: 28
% Final passive clauses: 10
% Main loop iterations started: 46
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 4
% Fw demodulations: 15
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Forward subsumptions: 60
% Fw demodulations to eq. taut.: 3
% 
% >>> Generating Inferences
% Binary resolution: 6
% Forward superposition: 58
% Backward superposition: 47
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 5
% SAT solver unit clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption resolution: 0.002 s
% other: 0.003 s ( own 0.0000 s ) 

