i: [1,31,0,6,3,1,0,0] t19_orders_1
a: [3,61,0,2,1,31]
s: 61
i: [1,30,0,2,1,1,0,0] t19_orders_1
a: [3,74,0,2,1,30]
s: 74
i: [1,29,0,2,1,1,0,0] t19_orders_1
a: [3,79,0,2,1,29]
s: 79
i: [1,32,0,6,3,0,0,1] d3_orders_1
s: 32
i: [1,33,0,6,3,0,0,1] d3_orders_1
s: 33
i: [1,35,0,6,3,0,0,1] d4_orders_1
s: 35
i: [1,36,0,6,3,0,0,1] d4_orders_1
s: 36
i: [1,37,0,6,3,0,0,1] d4_orders_1
s: 37
i: [1,39,0,6,3,0,0,1] d5_orders_1
s: 39
i: [1,40,0,6,3,0,0,1] d5_orders_1
s: 40
i: [1,41,0,6,3,0,0,1] d5_orders_1
s: 41
i: [1,42,0,6,3,0,0,1] d5_orders_1
s: 42
i: [1,44,0,6,3,0,0,255] d4_wellord1
s: 44
d: [2,81,1,4,2,1,40,44,74]
d: [2,82,1,2,1,2,21,81,79]
a: [3,86,1,2,1,82]
s: 86
i: [1,45,0,6,3,0,0,255] d4_wellord1
s: 45
i: [1,46,0,6,3,0,0,255] d4_wellord1
s: 46
d: [2,88,1,4,2,1,40,45,74]
d: [2,89,1,2,1,2,21,88,79]
a: [3,93,1,2,1,89]
s: 93
i: [1,47,0,6,3,0,0,255] d4_wellord1
s: 47
d: [2,95,1,4,2,1,40,46,74]
d: [2,96,1,2,1,2,21,95,79]
a: [3,100,1,2,1,96]
s: 100
i: [1,48,0,6,3,0,0,255] d4_wellord1
s: 48
i: [1,53,0,6,3,0,0,255] cc1_wellord1
s: 53
d: [2,102,1,4,2,1,40,47,74]
d: [2,103,1,2,1,2,21,102,79]
a: [3,107,1,2,1,103]
s: 107
i: [1,54,0,6,3,0,0,255] cc1_wellord1
s: 54
d: [2,109,1,4,2,1,40,48,74]
d: [2,110,1,2,1,2,21,109,79]
a: [3,114,1,2,1,110]
s: 114
i: [1,55,0,6,3,0,0,255] cc1_wellord1
s: 55
i: [1,56,0,6,3,0,0,255] cc1_wellord1
s: 56
i: [1,57,0,6,3,0,0,255] cc1_wellord1
s: 57
i: [1,34,0,8,4,0,0,1] d3_orders_1
s: 34
d: [2,116,2,0,0,4,52,79,86,61,93,34]
e: 116
d: [2,117,2,6,3,1,40,34,93]
d: [2,119,2,4,2,2,21,117,61]
d: [2,120,2,2,1,3,21,119,86]
d: [2,121,2,0,0,4,21,120,79]
e: 121
r: 61
a: [3,60,0,2,1,31]
s: 60
a: [3,65,0,2,1,31]
s: 65
i: [1,38,0,10,5,0,0,1] d4_orders_1
s: 38
d: [2,128,2,0,0,5,52,79,86,65,100,93,38]
e: 128
d: [2,129,2,8,4,1,40,38,93]
d: [2,131,2,6,3,2,21,129,100]
d: [2,132,2,4,2,3,21,131,65]
d: [2,133,2,2,1,4,21,132,86]
d: [2,134,2,0,0,5,21,133,79]
e: 134
r: 65
a: [3,64,0,2,1,31]
s: 64
a: [3,69,0,2,1,31]
s: 69
i: [1,43,0,12,6,0,0,1] d5_orders_1
s: 43
d: [2,142,2,0,0,6,52,79,86,69,100,107,93,43]
e: 142
d: [2,143,2,10,5,1,40,43,93]
d: [2,145,2,8,4,2,21,143,107]
d: [2,146,2,6,3,3,21,145,100]
d: [2,147,2,4,2,4,21,146,69]
d: [2,148,2,2,1,5,21,147,86]
d: [2,149,2,0,0,6,21,148,79]
e: 149
d: [2,123,2,4,2,2,21,117,86]
d: [2,124,2,2,1,3,21,123,79]
d: [2,136,2,4,2,3,21,131,86]
d: [2,137,2,2,1,4,21,136,79]
f: 149,137,124,103,96,89,82,29,30,31
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t19_orders_1
% SZS output start Proof for orders_1__t19_orders_1
1. ! [X0] : (v1_relat_1(X0) => (v2_wellord1(X0) => (v3_orders_1(X0) & v2_orders_1(X0) & v1_orders_1(X0)))) [input t19_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => (v2_wellord1(X0) => (v3_orders_1(X0) & v2_orders_1(X0) & v1_orders_1(X0)))) [negated conjecture 1]
5. ! [X0] : (v1_relat_1(X0) => (v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0)))) [input d3_orders_1]
6. ! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d4_orders_1]
7. ! [X0] : (v1_relat_1(X0) => (v2_wellord1(X0) <=> (v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d4_wellord1]
8. ! [X0] : (v1_relat_1(X0) => (v3_orders_1(X0) <=> (v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d5_orders_1]
9. ? [X0] : (((~v3_orders_1(X0) | ~v2_orders_1(X0) | ~v1_orders_1(X0)) & v2_wellord1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
10. ? [X0] : ((~v3_orders_1(X0) | ~v2_orders_1(X0) | ~v1_orders_1(X0)) & v2_wellord1(X0) & v1_relat_1(X0)) [flattening 9]
11. ! [X0] : ((v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 5]
12. ! [X0] : ((v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 6]
13. ! [X0] : ((v3_orders_1(X0) <=> (v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 8]
14. ! [X0] : ((v2_wellord1(X0) <=> (v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 7]
19. ? [X0] : ((~v3_orders_1(X0) | ~v2_orders_1(X0) | ~v1_orders_1(X0)) & v2_wellord1(X0) & v1_relat_1(X0)) => ((~v3_orders_1(sK0) | ~v2_orders_1(sK0) | ~v1_orders_1(sK0)) & v2_wellord1(sK0) & v1_relat_1(sK0)) [choice axiom]
20. (~v3_orders_1(sK0) | ~v2_orders_1(sK0) | ~v1_orders_1(sK0)) & v2_wellord1(sK0) & v1_relat_1(sK0) [skolemisation 10,19]
21. ! [X0] : (((v1_orders_1(X0) | (~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 11]
22. ! [X0] : (((v1_orders_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 21]
23. ! [X0] : (((v2_orders_1(X0) | (~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 12]
24. ! [X0] : (((v2_orders_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 23]
25. ! [X0] : (((v3_orders_1(X0) | (~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v3_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 13]
26. ! [X0] : (((v3_orders_1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v3_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 25]
27. ! [X0] : (((v2_wellord1(X0) | (~v1_wellord1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_wellord1(X0))) | ~v1_relat_1(X0)) [nnf transformation 14]
28. ! [X0] : (((v2_wellord1(X0) | ~v1_wellord1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_wellord1(X0))) | ~v1_relat_1(X0)) [flattening 27]
29. v1_relat_1(sK0) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. v2_wellord1(sK0) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~v3_orders_1(sK0) | ~v2_orders_1(sK0) | ~v1_orders_1(sK0) [cnf transformation 20] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. ~v8_relat_2(X0) | v1_orders_1(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
38. ~v8_relat_2(X0) | ~v4_relat_2(X0) | v2_orders_1(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. ~v8_relat_2(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | v3_orders_1(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. ~v2_wellord1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
45. ~v2_wellord1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
46. ~v2_wellord1(X0) | v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
47. ~v2_wellord1(X0) | v6_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
59. 1 <=> v1_orders_1(sK0) [avatar definition]
63. 2 <=> v2_orders_1(sK0) [avatar definition]
67. 3 <=> v3_orders_1(sK0) [avatar definition]
69. ~v3_orders_1(sK0) <- (~3) [avatar component clause 67] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
70. ~1 | ~2 | ~3 [avatar split clause 31,67,63,59]
72. 4 <=> v2_wellord1(sK0) [avatar definition]
74. v2_wellord1(sK0) <- (4) [avatar component clause 72] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. 4 [avatar split clause 30,72]
77. 5 <=> v1_relat_1(sK0) [avatar definition]
79. v1_relat_1(sK0) <- (5) [avatar component clause 77] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. 5 [avatar split clause 29,77]
81. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (4) [resolution 44,74] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. v1_relat_2(sK0) <- (4, 5) [subsumption resolution 81,79] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
84. 6 <=> v1_relat_2(sK0) [avatar definition]
86. v1_relat_2(sK0) <- (6) [avatar component clause 84] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
87. 6 | ~4 | ~5 [avatar split clause 82,77,72,84]
88. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (4) [resolution 45,74] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. v8_relat_2(sK0) <- (4, 5) [subsumption resolution 88,79] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. 7 <=> v8_relat_2(sK0) [avatar definition]
93. v8_relat_2(sK0) <- (7) [avatar component clause 91] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
94. 7 | ~4 | ~5 [avatar split clause 89,77,72,91]
95. v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (4) [resolution 46,74] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. v4_relat_2(sK0) <- (4, 5) [subsumption resolution 95,79] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. 8 <=> v4_relat_2(sK0) [avatar definition]
100. v4_relat_2(sK0) <- (8) [avatar component clause 98] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. 8 | ~4 | ~5 [avatar split clause 96,77,72,98]
102. v6_relat_2(sK0) | ~v1_relat_1(sK0) <- (4) [resolution 47,74] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
103. v6_relat_2(sK0) <- (4, 5) [subsumption resolution 102,79] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 9 <=> v6_relat_2(sK0) [avatar definition]
107. v6_relat_2(sK0) <- (9) [avatar component clause 105] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
108. 9 | ~4 | ~5 [avatar split clause 103,77,72,105]
117. v1_orders_1(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (7) [resolution 34,93] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
123. v1_orders_1(sK0) | ~v1_relat_1(sK0) <- (6, 7) [subsumption resolution 117,86] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
124. v1_orders_1(sK0) <- (5, 6, 7) [subsumption resolution 123,79] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
125. 1 | ~5 | ~6 | ~7 [avatar split clause 124,91,84,77,59]
129. ~v4_relat_2(sK0) | v2_orders_1(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (7) [resolution 38,93] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
131. v2_orders_1(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (7, 8) [subsumption resolution 129,100] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
136. v2_orders_1(sK0) | ~v1_relat_1(sK0) <- (6, 7, 8) [subsumption resolution 131,86] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
137. v2_orders_1(sK0) <- (5, 6, 7, 8) [subsumption resolution 136,79] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
138. 2 | ~5 | ~6 | ~7 | ~8 [avatar split clause 137,98,91,84,77,63]
143. ~v6_relat_2(sK0) | ~v4_relat_2(sK0) | v3_orders_1(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (7) [resolution 43,93] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
145. ~v4_relat_2(sK0) | v3_orders_1(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (7, 9) [subsumption resolution 143,107] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
146. v3_orders_1(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (7, 8, 9) [subsumption resolution 145,100] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
147. ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (~3, 7, 8, 9) [subsumption resolution 146,69] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
148. ~v1_relat_1(sK0) <- (~3, 6, 7, 8, 9) [subsumption resolution 147,86] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
149. $false <- (~3, 5, 6, 7, 8, 9) [subsumption resolution 148,79] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
150. 3 | ~5 | ~6 | ~7 | ~8 | ~9 [avatar contradiction clause 149]
151. $false [avatar sat refutation 70,75,80,87,94,101,108,125,138,150] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:112,thDist:-112}
% SZS output end Proof for orders_1__t19_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 7
% Introduced skolems: 1
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 80
% Active clauses: 29
% Passive clauses: 36
% Final active clauses: 27
% Final passive clauses: 2
% Main loop iterations started: 34
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 21
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 3
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 3
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 2
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

