i: [1,108,0,2,1,1,0,0] t10_tdlat_3
a: [3,158,0,2,1,108]
s: 158
i: [1,107,0,2,1,1,0,0] t10_tdlat_3
a: [3,163,0,2,1,107]
s: 163
i: [1,106,0,2,1,1,0,0] t10_tdlat_3
a: [3,168,0,2,1,106]
s: 168
i: [1,109,0,8,1,1,0,0] t10_tdlat_3
a: [3,153,0,8,1,109]
s: 153
i: [1,110,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,173,0,2,1,110]
s: 173
i: [1,111,0,2,1,0,0,1] fc1_xboole_0
a: [3,178,0,2,1,111]
s: 178
i: [1,147,0,2,1,0,0,2] rc2_xboole_0
a: [3,188,0,2,1,147]
s: 188
i: [1,148,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,193,0,2,1,148]
s: 193
i: [1,149,0,2,1,0,0,2] rc1_xboole_0
a: [3,198,0,2,1,149]
s: 198
i: [1,113,0,3,1,0,0,2] fc1_subset_1
s: 113
i: [1,140,0,3,1,0,0,2] rc2_subset_1
s: 140
i: [1,143,0,3,1,0,0,255] reflexivity_r1_tarski
s: 143
i: [1,112,0,3,1,0,0,255] d2_xboole_0
a: [3,183,0,3,1,112]
s: 183
r: 173
i: [1,138,0,4,1,0,0,2] existence_m1_subset_1
s: 138
i: [1,142,0,4,1,0,0,255] rc3_subset_1
s: 142
i: [1,144,0,4,1,0,0,1] fc3_xboole_0
s: 144
i: [1,114,0,5,1,0,0,1] redefinition_k9_setfam_1
s: 114
r: 113
d: [2,201,0,3,1,0,23,113,114]
s: 201
i: [1,119,0,5,2,0,0,2] rc1_subset_1
s: 119
i: [1,128,0,5,2,0,0,1] t6_boole
s: 128
i: [1,135,0,5,2,0,0,255] t7_boole
s: 135
d: [2,204,1,10,2,1,44,153,128]
s: 204
i: [1,139,0,5,1,0,0,2] rc2_subset_1
s: 139
i: [1,141,0,5,1,0,0,255] rc3_subset_1
s: 141
d: [2,205,0,5,1,0,22,139,114]
s: 205
d: [2,206,0,5,1,0,22,141,114]
s: 206
i: [1,115,0,6,1,0,0,1] dt_k9_setfam_1
s: 115
i: [1,121,0,6,2,0,0,255] rc4_subset_1
s: 121
i: [1,131,0,6,2,0,0,255] t1_subset
s: 131
i: [1,146,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 146
d: [2,207,0,6,1,0,22,115,114]
d: [2,208,0,6,1,0,22,207,114]
s: 208
i: [1,118,0,7,2,0,0,2] rc1_subset_1
s: 118
i: [1,120,0,7,2,0,0,255] rc4_subset_1
s: 120
i: [1,133,0,7,2,0,0,255] t3_subset
s: 133
i: [1,134,0,7,3,0,0,1] t8_boole
s: 134
i: [1,145,0,7,1,0,0,1] commutativity_k2_tarski
s: 145
d: [2,213,2,10,2,1,44,204,145]
s: 213
d: [2,209,0,7,2,0,22,118,114]
s: 209
d: [2,215,2,10,2,1,44,204,145]
s: 215
d: [2,210,0,7,2,0,22,120,114]
s: 210
d: [2,211,0,7,2,0,22,133,114]
s: 211
d: [2,202,1,7,3,0,44,128,128]
s: 202
i: [1,127,0,8,2,0,0,255] rc5_tops_1
s: 127
i: [1,129,0,8,3,0,0,2] cc1_subset_1
s: 129
i: [1,132,0,8,3,0,0,255] t2_subset
s: 132
d: [2,218,0,8,2,0,22,127,114]
s: 218
d: [2,225,1,8,2,0,40,132,205]
d: [2,231,1,5,1,0,21,225,201]
s: 231
d: [2,219,0,8,3,0,22,129,114]
s: 219
d: [2,226,1,8,2,0,40,132,206]
d: [2,232,1,5,1,0,21,226,201]
s: 232
i: [1,122,0,9,2,0,0,2] dt_u1_pre_topc
s: 122
i: [1,130,0,9,3,0,0,255] cc4_subset_1
s: 130
d: [2,245,1,5,2,0,40,219,206]
s: 245
i: [1,137,0,9,3,0,0,255] t5_subset
s: 137
d: [2,235,2,5,1,0,40,231,146]
s: 235
d: [2,249,0,9,2,0,22,122,114]
d: [2,250,0,9,2,0,22,249,114]
s: 250
d: [2,251,0,9,3,0,22,130,114]
s: 251
d: [2,246,2,5,1,0,40,232,146]
s: 246
d: [2,252,0,9,3,0,22,137,114]
s: 252
d: [2,224,1,6,2,0,40,132,138]
s: 224
i: [1,123,0,10,3,0,0,1] d1_tdlat_3
s: 123
i: [1,124,0,10,3,0,0,1] d1_tdlat_3
s: 124
d: [2,221,1,10,2,0,40,132,208]
d: [2,228,1,6,1,0,21,221,201]
s: 228
i: [1,136,0,10,3,0,0,255] t4_subset
s: 136
d: [2,243,1,6,2,0,40,219,138]
s: 243
d: [2,301,0,10,3,0,22,136,114]
s: 301
i: [1,116,0,11,4,0,0,255] cc3_subset_1
s: 116
d: [2,263,1,6,2,0,40,251,205]
s: 263
i: [1,117,0,11,4,0,0,255] cc2_subset_1
s: 117
d: [2,272,1,6,2,0,40,252,205]
s: 272
i: [1,125,0,11,3,0,0,1] d2_tdlat_3
s: 125
i: [1,126,0,11,3,0,0,1] d2_tdlat_3
s: 126
d: [2,317,1,10,3,1,44,153,125]
d: [2,320,1,8,2,2,21,317,158]
d: [2,321,1,6,1,3,21,320,168]
a: [3,325,1,6,1,321]
s: 325
d: [2,334,2,0,0,3,52,168,163,325,123]
e: 334
d: [2,335,2,9,3,1,44,325,123]
d: [2,336,2,4,2,1,19,335]
d: [2,338,2,2,1,2,21,336,163]
d: [2,339,2,0,0,3,21,338,168]
e: 339
d: [2,318,2,12,4,1,44,204,125]
d: [2,327,2,10,3,2,21,318,178]
d: [2,328,2,8,2,3,21,327,158]
d: [2,329,2,6,1,4,21,328,168]
f: 339,329,111,106,107,108,109
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tdlat_3__t10_tdlat_3
% SZS output start Proof for tdlat_3__t10_tdlat_3
1. ! [X0] : (l1_pre_topc(X0) => ((v2_tdlat_3(X0) & v1_tdlat_3(X0)) => k9_setfam_1(u1_struct_0(X0)) = k2_tarski(k1_xboole_0,u1_struct_0(X0)))) [input t10_tdlat_3]
2. ~! [X0] : (l1_pre_topc(X0) => ((v2_tdlat_3(X0) & v1_tdlat_3(X0)) => k9_setfam_1(u1_struct_0(X0)) = k2_tarski(k1_xboole_0,u1_struct_0(X0)))) [negated conjecture 1]
10. ! [X0] : (l1_pre_topc(X0) => (v1_tdlat_3(X0) <=> k9_setfam_1(u1_struct_0(X0)) = u1_pre_topc(X0))) [input d1_tdlat_3]
11. ! [X0] : (l1_pre_topc(X0) => (v2_tdlat_3(X0) <=> k2_tarski(k1_xboole_0,u1_struct_0(X0)) = u1_pre_topc(X0))) [input d2_tdlat_3]
27. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
43. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
59. ? [X0] : ((k9_setfam_1(u1_struct_0(X0)) != k2_tarski(k1_xboole_0,u1_struct_0(X0)) & (v2_tdlat_3(X0) & v1_tdlat_3(X0))) & l1_pre_topc(X0)) [ennf transformation 2]
60. ? [X0] : (k9_setfam_1(u1_struct_0(X0)) != k2_tarski(k1_xboole_0,u1_struct_0(X0)) & v2_tdlat_3(X0) & v1_tdlat_3(X0) & l1_pre_topc(X0)) [flattening 59]
68. ! [X0] : ((v1_tdlat_3(X0) <=> k9_setfam_1(u1_struct_0(X0)) = u1_pre_topc(X0)) | ~l1_pre_topc(X0)) [ennf transformation 10]
69. ! [X0] : ((v2_tdlat_3(X0) <=> k2_tarski(k1_xboole_0,u1_struct_0(X0)) = u1_pre_topc(X0)) | ~l1_pre_topc(X0)) [ennf transformation 11]
71. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 43]
84. ? [X0] : (k9_setfam_1(u1_struct_0(X0)) != k2_tarski(k1_xboole_0,u1_struct_0(X0)) & v2_tdlat_3(X0) & v1_tdlat_3(X0) & l1_pre_topc(X0)) => (k9_setfam_1(u1_struct_0(sK0)) != k2_tarski(k1_xboole_0,u1_struct_0(sK0)) & v2_tdlat_3(sK0) & v1_tdlat_3(sK0) & l1_pre_topc(sK0)) [choice axiom]
85. k9_setfam_1(u1_struct_0(sK0)) != k2_tarski(k1_xboole_0,u1_struct_0(sK0)) & v2_tdlat_3(sK0) & v1_tdlat_3(sK0) & l1_pre_topc(sK0) [skolemisation 60,84]
90. ! [X0] : (((v1_tdlat_3(X0) | k9_setfam_1(u1_struct_0(X0)) != u1_pre_topc(X0)) & (k9_setfam_1(u1_struct_0(X0)) = u1_pre_topc(X0) | ~v1_tdlat_3(X0))) | ~l1_pre_topc(X0)) [nnf transformation 68]
91. ! [X0] : (((v2_tdlat_3(X0) | k2_tarski(k1_xboole_0,u1_struct_0(X0)) != u1_pre_topc(X0)) & (k2_tarski(k1_xboole_0,u1_struct_0(X0)) = u1_pre_topc(X0) | ~v2_tdlat_3(X0))) | ~l1_pre_topc(X0)) [nnf transformation 69]
106. l1_pre_topc(sK0) [cnf transformation 85] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
107. v1_tdlat_3(sK0) [cnf transformation 85] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
108. v2_tdlat_3(sK0) [cnf transformation 85] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
109. k9_setfam_1(u1_struct_0(sK0)) != k2_tarski(k1_xboole_0,u1_struct_0(sK0)) [cnf transformation 85] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
111. v1_xboole_0(k1_xboole_0) [cnf transformation 27] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
123. k9_setfam_1(u1_struct_0(X0)) = u1_pre_topc(X0) | ~v1_tdlat_3(X0) | ~l1_pre_topc(X0) [cnf transformation 90] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
125. k2_tarski(k1_xboole_0,u1_struct_0(X0)) = u1_pre_topc(X0) | ~v2_tdlat_3(X0) | ~l1_pre_topc(X0) [cnf transformation 91] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
128. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 71] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
151. 1 <=> k9_setfam_1(u1_struct_0(sK0)) = k2_tarski(k1_xboole_0,u1_struct_0(sK0)) [avatar definition]
153. k9_setfam_1(u1_struct_0(sK0)) != k2_tarski(k1_xboole_0,u1_struct_0(sK0)) <- (~1) [avatar component clause 151] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
154. ~1 [avatar split clause 109,151]
156. 2 <=> v2_tdlat_3(sK0) [avatar definition]
158. v2_tdlat_3(sK0) <- (2) [avatar component clause 156] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
159. 2 [avatar split clause 108,156]
161. 3 <=> v1_tdlat_3(sK0) [avatar definition]
163. v1_tdlat_3(sK0) <- (3) [avatar component clause 161] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
164. 3 [avatar split clause 107,161]
166. 4 <=> l1_pre_topc(sK0) [avatar definition]
168. l1_pre_topc(sK0) <- (4) [avatar component clause 166] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
169. 4 [avatar split clause 106,166]
176. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
178. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 176] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
179. 6 [avatar split clause 111,176]
204. k9_setfam_1(u1_struct_0(sK0)) != k2_tarski(X0,u1_struct_0(sK0)) | ~v1_xboole_0(X0) <- (~1) [superposition 153,128] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
318. k9_setfam_1(u1_struct_0(sK0)) != u1_pre_topc(sK0) | ~v1_xboole_0(k1_xboole_0) | ~v2_tdlat_3(sK0) | ~l1_pre_topc(sK0) <- (~1) [superposition 204,125] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
323. 11 <=> k9_setfam_1(u1_struct_0(sK0)) = u1_pre_topc(sK0) [avatar definition]
325. k9_setfam_1(u1_struct_0(sK0)) != u1_pre_topc(sK0) <- (~11) [avatar component clause 323] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
327. k9_setfam_1(u1_struct_0(sK0)) != u1_pre_topc(sK0) | ~v2_tdlat_3(sK0) | ~l1_pre_topc(sK0) <- (~1, 6) [subsumption resolution 318,178] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
328. k9_setfam_1(u1_struct_0(sK0)) != u1_pre_topc(sK0) | ~l1_pre_topc(sK0) <- (~1, 2, 6) [subsumption resolution 327,158] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
329. k9_setfam_1(u1_struct_0(sK0)) != u1_pre_topc(sK0) <- (~1, 2, 4, 6) [subsumption resolution 328,168] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
330. ~11 | 1 | ~2 | ~4 | ~6 [avatar split clause 329,176,166,156,151,323]
335. u1_pre_topc(sK0) != u1_pre_topc(sK0) | ~v1_tdlat_3(sK0) | ~l1_pre_topc(sK0) <- (~11) [superposition 325,123] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
336. ~v1_tdlat_3(sK0) | ~l1_pre_topc(sK0) <- (~11) [trivial inequality removal 335] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
338. ~l1_pre_topc(sK0) <- (3, ~11) [subsumption resolution 336,163] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
339. $false <- (3, 4, ~11) [subsumption resolution 338,168] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
340. ~3 | ~4 | 11 [avatar contradiction clause 339]
341. $false [avatar sat refutation 154,159,164,169,179,330,340] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:55,thDist:-55}
% SZS output end Proof for tdlat_3__t10_tdlat_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 44
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 44
% Generated clauses: 188
% Active clauses: 56
% Passive clauses: 120
% Final active clauses: 54
% Final passive clauses: 48
% Main loop iterations started: 72
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Trivial inequalities: 1
% Fw subsumption resolutions: 15
% Fw demodulations: 16
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 58
% Unit resulting resolution: 1
% Forward superposition: 10
% Backward superposition: 25
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 10
% 
% Memory used [KB]: 9466
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% other: 0.004 s ( own 0.002 s ) 

