i: [1,165,0,3,1,1,0,0] l25_binop_2
a: [3,223,0,3,1,165]
s: 223
i: [1,166,0,2,1,0,0,2] fc1_numbers
a: [3,228,0,2,1,166]
s: 228
i: [1,173,0,2,1,0,0,3] fc6_ordinal1
a: [3,258,0,2,1,173]
s: 258
i: [1,172,0,2,1,0,0,3] fc6_ordinal1
a: [3,263,0,2,1,172]
s: 263
i: [1,204,0,2,1,0,0,4] rc4_ordinal1
a: [3,268,0,2,1,204]
s: 268
i: [1,205,0,2,1,0,0,3] rc1_ordinal1
a: [3,273,0,2,1,205]
s: 273
i: [1,209,0,2,1,0,0,255] rc2_ordinal1
a: [3,278,0,2,1,209]
s: 278
i: [1,208,0,2,1,0,0,255] rc2_ordinal1
a: [3,283,0,2,1,208]
s: 283
i: [1,207,0,2,1,0,0,255] rc2_ordinal1
a: [3,288,0,2,1,207]
s: 288
i: [1,206,0,2,1,0,0,255] rc2_ordinal1
a: [3,293,0,2,1,206]
s: 293
i: [1,211,0,2,1,0,0,4] rc5_ordinal1
a: [3,298,0,2,1,211]
s: 298
i: [1,210,0,2,1,0,0,4] rc5_ordinal1
a: [3,303,0,2,1,210]
s: 303
i: [1,212,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,308,0,2,1,212]
s: 308
i: [1,213,0,2,1,0,0,255] rc4_xreal_0
a: [3,313,0,2,1,213]
s: 313
i: [1,214,0,2,1,0,0,255] rc3_xreal_0
a: [3,318,0,2,1,214]
s: 318
i: [1,219,0,2,1,0,0,2] rc2_nat_1
a: [3,333,0,2,1,219]
s: 333
i: [1,218,0,2,1,0,0,2] rc2_nat_1
a: [3,338,0,2,1,218]
s: 338
i: [1,198,0,3,1,0,0,255] reflexivity_r1_tarski
s: 198
i: [1,167,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,233,0,3,1,167]
s: 233
i: [1,168,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,238,0,3,1,168]
s: 238
r: 263
r: 258
d: [2,345,0,2,1,2,23,263,238]
a: [3,350,0,2,1,345]
s: 350
d: [2,346,0,2,1,2,23,258,238]
a: [3,355,0,2,1,346]
s: 355
i: [1,169,0,3,1,0,0,2] t1_numerals
a: [3,243,0,3,1,169]
s: 243
i: [1,215,0,3,1,0,0,255] rc1_rat_1
a: [3,323,0,3,1,215]
s: 323
i: [1,216,0,3,1,0,0,255] rc1_int_1
a: [3,328,0,3,1,216]
s: 328
d: [2,357,0,3,1,2,22,243,238]
d: [2,358,0,3,1,3,22,357,233]
a: [3,362,0,3,1,358]
s: 362
i: [1,175,0,4,2,0,0,3] cc6_ordinal1
s: 175
i: [1,177,0,4,2,0,0,255] cc3_nat_1
s: 177
d: [2,364,1,2,1,1,40,175,268]
a: [3,369,1,2,1,364]
s: 369
i: [1,178,0,4,2,0,0,3] cc1_nat_1
s: 178
i: [1,180,0,4,2,0,0,255] cc9_ordinal1
s: 180
d: [2,365,1,2,1,1,40,175,298]
a: [3,374,1,2,1,365]
s: 374
i: [1,181,0,4,2,0,0,4] cc7_ordinal1
s: 181
d: [2,376,1,2,1,1,40,177,268]
a: [3,381,1,2,1,376]
s: 381
i: [1,182,0,4,2,0,0,3] cc3_ordinal1
s: 182
i: [1,184,0,4,2,0,0,255] cc1_ordinal1
s: 184
d: [2,377,1,2,1,1,40,177,298]
a: [3,386,1,2,1,377]
s: 386
i: [1,185,0,4,2,0,0,255] cc1_ordinal1
s: 185
d: [2,388,1,2,1,1,40,180,313]
a: [3,392,1,2,1,388]
s: 392
i: [1,197,0,4,1,0,0,2] existence_m1_subset_1
s: 197
i: [1,170,0,4,1,0,0,1] dt_k5_numbers
a: [3,248,0,4,1,170]
s: 248
d: [2,394,1,2,1,1,40,181,313]
a: [3,398,1,2,1,394]
s: 398
i: [1,171,0,4,1,0,0,1] dt_k6_numbers
a: [3,253,0,4,1,171]
s: 253
d: [2,400,1,2,1,1,40,182,313]
a: [3,404,1,2,1,400]
s: 404
i: [1,217,0,4,1,0,0,2] rc2_nat_1
a: [3,343,0,4,1,217]
s: 343
i: [1,183,0,5,2,0,0,2] t6_boole
s: 183
d: [2,406,1,2,1,1,40,184,355]
a: [3,415,1,2,1,406]
s: 415
i: [1,187,0,5,2,0,0,255] cc2_nat_1
s: 187
d: [2,408,1,2,1,1,40,184,273]
a: [3,425,1,2,1,408]
s: 425
i: [1,188,0,5,2,0,0,3] cc8_ordinal1
s: 188
i: [1,194,0,5,2,0,0,1] t7_boole
s: 194
d: [2,411,1,2,1,1,40,184,333]
a: [3,435,1,2,1,411]
s: 435
d: [2,488,0,5,2,1,22,183,233]
s: 488
d: [2,505,1,5,2,2,44,223,488]
s: 505
d: [2,501,0,5,2,1,22,188,238]
s: 501
i: [1,189,0,6,3,0,0,255] cc2_ordinal1
s: 189
d: [2,437,1,2,1,1,40,185,355]
a: [3,446,1,2,1,437]
s: 446
i: [1,190,0,6,2,0,0,1] t1_subset
s: 190
d: [2,439,1,2,1,1,40,185,273]
a: [3,456,1,2,1,439]
s: 456
i: [1,199,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 199
i: [1,186,0,7,3,0,0,3] cc5_ordinal1
s: 186
d: [2,442,1,2,1,1,40,185,333]
a: [3,466,1,2,1,442]
s: 466
i: [1,192,0,7,2,0,0,255] t3_subset
s: 192
d: [2,489,1,2,1,1,40,187,362]
a: [3,494,1,2,1,489]
s: 494
i: [1,193,0,7,3,0,0,2] t8_boole
s: 193
i: [1,174,0,8,3,0,0,255] cc10_ordinal1
s: 174
d: [2,506,1,2,1,2,40,501,362]
a: [3,511,1,2,1,506]
s: 511
i: [1,191,0,8,3,0,0,1] t2_subset
s: 191
d: [2,572,1,0,0,3,52,350,223,362,191]
e: 572
d: [2,574,1,5,2,1,40,191,362]
d: [2,581,1,3,1,2,21,574,350]
d: [2,582,1,0,0,3,21,581,223]
e: 582
f: 582,358,345,172,169,168,167,165
% Refutation found. Thanks to Tanya!
% SZS status Theorem for binop_2__l25_binop_2
% SZS output start Proof for binop_2__l25_binop_2
1. r2_hidden(k6_numbers,k5_numbers) [input l25_binop_2]
2. ~r2_hidden(k6_numbers,k5_numbers) [negated conjecture 1]
52. v3_ordinal1(k4_ordinal1) & ~v1_xboole_0(k4_ordinal1) [input fc6_ordinal1]
54. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
55. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
62. k5_numbers = k4_ordinal1 [input redefinition_k5_numbers]
63. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
68. ~r2_hidden(k6_numbers,k5_numbers) [flattening 2]
70. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 55]
125. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 70]
126. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 125]
165. ~r2_hidden(k6_numbers,k5_numbers) [cnf transformation 68] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
167. k6_numbers = k1_xboole_0 [cnf transformation 63] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
168. k5_numbers = k4_ordinal1 [cnf transformation 62] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
169. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 54] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
172. ~v1_xboole_0(k4_ordinal1) [cnf transformation 52] {a:0,w:2,wCS:800,sine:3,thAx:0,allAx:1,thDist:-1}
191. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 126] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
221. 1 <=> r2_hidden(k6_numbers,k5_numbers) [avatar definition]
223. ~r2_hidden(k6_numbers,k5_numbers) <- (~1) [avatar component clause 221] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
224. ~1 [avatar split clause 165,221]
231. 3 <=> k6_numbers = k1_xboole_0 [avatar definition]
233. k6_numbers = k1_xboole_0 <- (3) [avatar component clause 231] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
234. 3 [avatar split clause 167,231]
236. 4 <=> k5_numbers = k4_ordinal1 [avatar definition]
238. k5_numbers = k4_ordinal1 <- (4) [avatar component clause 236] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
239. 4 [avatar split clause 168,236]
241. 5 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
243. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (5) [avatar component clause 241] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
244. 5 [avatar split clause 169,241]
261. 9 <=> v1_xboole_0(k4_ordinal1) [avatar definition]
263. ~v1_xboole_0(k4_ordinal1) <- (~9) [avatar component clause 261] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
264. ~9 [avatar split clause 172,261]
345. ~v1_xboole_0(k5_numbers) <- (4, ~9) [backward demodulation 263,238] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
348. 26 <=> v1_xboole_0(k5_numbers) [avatar definition]
350. ~v1_xboole_0(k5_numbers) <- (~26) [avatar component clause 348] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
351. ~26 | ~4 | 9 [avatar split clause 345,261,236,348]
357. m1_subset_1(k1_xboole_0,k5_numbers) <- (4, 5) [forward demodulation 243,238] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
358. m1_subset_1(k6_numbers,k5_numbers) <- (3, 4, 5) [forward demodulation 357,233] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
360. 28 <=> m1_subset_1(k6_numbers,k5_numbers) [avatar definition]
362. m1_subset_1(k6_numbers,k5_numbers) <- (28) [avatar component clause 360] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
363. 28 | ~3 | ~4 | ~5 [avatar split clause 358,241,236,231,360]
574. v1_xboole_0(k5_numbers) | r2_hidden(k6_numbers,k5_numbers) <- (28) [resolution 191,362] {a:1,w:5,wCS:2000,sine:1,thAx:0,allAx:2,thDist:-2}
581. r2_hidden(k6_numbers,k5_numbers) <- (~26, 28) [subsumption resolution 574,350] {a:1,w:3,wCS:1200,sine:1,thAx:0,allAx:2,thDist:-2}
582. $false <- (~1, ~26, 28) [subsumption resolution 581,223] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
583. 1 | 26 | ~28 [avatar contradiction clause 582]
586. $false [avatar sat refutation 224,234,239,244,264,351,363,583] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:247,thDist:-247}
% SZS output end Proof for binop_2__l25_binop_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 66
% >>> Preprocessing
% Introduced skolems: 12
% Pure predicates: 7
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 55
% Generated clauses: 179
% Active clauses: 64
% Passive clauses: 94
% Final active clauses: 62
% Final passive clauses: 26
% Main loop iterations started: 68
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 5
% Fw demodulations: 4
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 51
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 63
% SAT solver unit clauses: 26
% SAT solver binary clauses: 24
% 
% Memory used [KB]: 9722
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

