i: [1,40,0,2,1,1,0,0] t34_orders_1
a: [3,69,0,2,1,40]
s: 69
i: [1,39,0,2,1,1,0,0] t34_orders_1
a: [3,74,0,2,1,39]
s: 74
i: [1,41,0,4,1,1,0,0] t34_orders_1
a: [3,64,0,4,1,41]
s: 64
i: [1,57,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 57
i: [1,47,0,6,3,0,0,1] d4_orders_1
s: 47
i: [1,48,0,6,3,0,0,1] d4_orders_1
s: 48
d: [2,76,1,4,2,1,40,47,69]
d: [2,77,1,2,1,2,21,76,74]
a: [3,81,1,2,1,77]
s: 81
i: [1,49,0,6,3,0,0,1] d4_orders_1
s: 49
d: [2,83,1,4,2,1,40,48,69]
d: [2,84,1,2,1,2,21,83,74]
a: [3,88,1,2,1,84]
s: 88
i: [1,58,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 58
i: [1,43,0,8,3,0,0,2] d12_relat_2
s: 43
d: [2,90,1,4,2,1,40,49,69]
d: [2,91,1,2,1,2,21,90,74]
a: [3,95,1,2,1,91]
s: 95
i: [1,44,0,8,3,0,0,2] d12_relat_2
s: 44
i: [1,45,0,8,3,0,0,2] d16_relat_2
s: 45
i: [1,46,0,8,3,0,0,2] d16_relat_2
s: 46
i: [1,51,0,8,3,0,0,2] d9_relat_2
s: 51
i: [1,52,0,8,3,0,0,2] d9_relat_2
s: 52
i: [1,53,0,8,3,0,0,1] d7_orders_1
s: 53
i: [1,54,0,8,3,0,0,1] d7_orders_1
s: 54
i: [1,55,0,8,3,0,0,1] d7_orders_1
s: 55
i: [1,60,0,8,3,0,0,3] fc3_relat_1
s: 60
i: [1,42,0,10,2,0,0,255] d6_relat_1
s: 42
i: [1,50,0,10,5,0,0,1] d4_orders_1
s: 50
d: [2,107,1,11,4,0,44,60,42]
s: 107
i: [1,59,0,12,5,0,0,2] fc4_relat_2
s: 59
i: [1,56,0,14,5,0,0,1] d7_orders_1
s: 56
d: [2,116,1,18,6,0,40,56,51]
d: [2,117,1,16,5,0,18,116]
s: 117
d: [2,113,1,17,6,0,44,59,42]
s: 113
d: [2,118,2,16,6,0,40,117,45]
d: [2,119,2,14,5,0,18,118]
s: 119
d: [2,109,1,24,8,0,40,59,50]
d: [2,115,1,20,7,0,21,109,60]
s: 115
d: [2,125,2,23,8,0,44,115,42]
s: 125
d: [2,120,3,14,6,0,40,119,43]
d: [2,121,3,12,5,0,18,120]
s: 121
d: [2,127,4,0,0,5,52,95,88,81,74,64,121]
e: 127
d: [2,128,4,8,4,1,40,121,64]
d: [2,136,4,6,3,2,21,128,74]
d: [2,137,4,4,2,3,21,136,81]
d: [2,138,4,2,1,4,21,137,88]
d: [2,139,4,0,0,5,21,138,95]
e: 139
f: 139,91,84,77,39,40,41
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t34_orders_1
% SZS output start Proof for orders_1__t34_orders_1
1. ! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) => r2_orders_1(X0,k1_relat_1(X0)))) [input t34_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) => r2_orders_1(X0,k1_relat_1(X0)))) [negated conjecture 1]
4. ! [X0] : (v1_relat_1(X0) => (v4_relat_2(X0) <=> r4_relat_2(X0,k1_relat_1(X0)))) [input d12_relat_2]
5. ! [X0] : (v1_relat_1(X0) => (v8_relat_2(X0) <=> r8_relat_2(X0,k1_relat_1(X0)))) [input d16_relat_2]
6. ! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d4_orders_1]
8. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r2_orders_1(X0,X1) <=> (r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)))) [input d7_orders_1]
9. ! [X0] : (v1_relat_1(X0) => (v1_relat_2(X0) <=> r1_relat_2(X0,k1_relat_1(X0)))) [input d9_relat_2]
18. ? [X0] : ((~r2_orders_1(X0,k1_relat_1(X0)) & v2_orders_1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
19. ? [X0] : (~r2_orders_1(X0,k1_relat_1(X0)) & v2_orders_1(X0) & v1_relat_1(X0)) [flattening 18]
21. ! [X0] : ((v4_relat_2(X0) <=> r4_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
22. ! [X0] : ((v8_relat_2(X0) <=> r8_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 5]
23. ! [X0] : ((v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 6]
24. ! [X0] : ((v1_relat_2(X0) <=> r1_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 9]
25. ! [X0] : (! [X1] : (r2_orders_1(X0,X1) <=> (r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 8]
30. ? [X0] : (~r2_orders_1(X0,k1_relat_1(X0)) & v2_orders_1(X0) & v1_relat_1(X0)) => (~r2_orders_1(sK0,k1_relat_1(sK0)) & v2_orders_1(sK0) & v1_relat_1(sK0)) [choice axiom]
31. ~r2_orders_1(sK0,k1_relat_1(sK0)) & v2_orders_1(sK0) & v1_relat_1(sK0) [skolemisation 19,30]
32. ! [X0] : (((v4_relat_2(X0) | ~r4_relat_2(X0,k1_relat_1(X0))) & (r4_relat_2(X0,k1_relat_1(X0)) | ~v4_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 21]
33. ! [X0] : (((v8_relat_2(X0) | ~r8_relat_2(X0,k1_relat_1(X0))) & (r8_relat_2(X0,k1_relat_1(X0)) | ~v8_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 22]
34. ! [X0] : (((v2_orders_1(X0) | (~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 23]
35. ! [X0] : (((v2_orders_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 34]
36. ! [X0] : (((v1_relat_2(X0) | ~r1_relat_2(X0,k1_relat_1(X0))) & (r1_relat_2(X0,k1_relat_1(X0)) | ~v1_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 24]
37. ! [X0] : (! [X1] : ((r2_orders_1(X0,X1) | (~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1))) & ((r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r2_orders_1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 25]
38. ! [X0] : (! [X1] : ((r2_orders_1(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1)) & ((r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r2_orders_1(X0,X1))) | ~v1_relat_1(X0)) [flattening 37]
39. v1_relat_1(sK0) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. v2_orders_1(sK0) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
41. ~r2_orders_1(sK0,k1_relat_1(sK0)) [cnf transformation 31] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. r4_relat_2(X0,k1_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 32] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
45. r8_relat_2(X0,k1_relat_1(X0)) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
47. ~v2_orders_1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
48. ~v2_orders_1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. ~v2_orders_1(X0) | v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. r1_relat_2(X0,k1_relat_1(X0)) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
56. ~r1_relat_2(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | r2_orders_1(X0,X1) | ~v1_relat_1(X0) [cnf transformation 38] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. 1 <=> r2_orders_1(sK0,k1_relat_1(sK0)) [avatar definition]
64. ~r2_orders_1(sK0,k1_relat_1(sK0)) <- (~1) [avatar component clause 62] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. ~1 [avatar split clause 41,62]
67. 2 <=> v2_orders_1(sK0) [avatar definition]
69. v2_orders_1(sK0) <- (2) [avatar component clause 67] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
70. 2 [avatar split clause 40,67]
72. 3 <=> v1_relat_1(sK0) [avatar definition]
74. v1_relat_1(sK0) <- (3) [avatar component clause 72] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. 3 [avatar split clause 39,72]
76. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 47,69] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
77. v1_relat_2(sK0) <- (2, 3) [subsumption resolution 76,74] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
79. 4 <=> v1_relat_2(sK0) [avatar definition]
81. v1_relat_2(sK0) <- (4) [avatar component clause 79] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. 4 | ~2 | ~3 [avatar split clause 77,72,67,79]
83. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 48,69] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
84. v8_relat_2(sK0) <- (2, 3) [subsumption resolution 83,74] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
86. 5 <=> v8_relat_2(sK0) [avatar definition]
88. v8_relat_2(sK0) <- (5) [avatar component clause 86] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. 5 | ~2 | ~3 [avatar split clause 84,72,67,86]
90. v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 49,69] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
91. v4_relat_2(sK0) <- (2, 3) [subsumption resolution 90,74] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
93. 6 <=> v4_relat_2(sK0) [avatar definition]
95. v4_relat_2(sK0) <- (6) [avatar component clause 93] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. 6 | ~2 | ~3 [avatar split clause 91,72,67,93]
116. ~r4_relat_2(X0,k1_relat_1(X0)) | ~r8_relat_2(X0,k1_relat_1(X0)) | r2_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [resolution 56,51] {a:1,w:18,wCS:7200,sine:1,thAx:0,allAx:2,thDist:-2}
117. ~r8_relat_2(X0,k1_relat_1(X0)) | ~r4_relat_2(X0,k1_relat_1(X0)) | r2_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) [duplicate literal removal 116] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
118. ~r4_relat_2(X0,k1_relat_1(X0)) | r2_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [resolution 117,45] {a:2,w:16,wCS:6400,sine:1,thAx:0,allAx:3,thDist:-3}
119. ~r4_relat_2(X0,k1_relat_1(X0)) | r2_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) | ~v8_relat_2(X0) [duplicate literal removal 118] {a:2,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
120. r2_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) | ~v8_relat_2(X0) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [resolution 119,43] {a:3,w:14,wCS:5600,sine:1,thAx:0,allAx:4,thDist:-4}
121. r2_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) | ~v8_relat_2(X0) | ~v4_relat_2(X0) [duplicate literal removal 120] {a:3,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
128. ~v1_relat_1(sK0) | ~v1_relat_2(sK0) | ~v8_relat_2(sK0) | ~v4_relat_2(sK0) <- (~1) [resolution 121,64] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
136. ~v1_relat_2(sK0) | ~v8_relat_2(sK0) | ~v4_relat_2(sK0) <- (~1, 3) [subsumption resolution 128,74] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
137. ~v8_relat_2(sK0) | ~v4_relat_2(sK0) <- (~1, 3, 4) [subsumption resolution 136,81] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
138. ~v4_relat_2(sK0) <- (~1, 3, 4, 5) [subsumption resolution 137,88] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
139. $false <- (~1, 3, 4, 5, 6) [subsumption resolution 138,95] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
140. 1 | ~3 | ~4 | ~5 | ~6 [avatar contradiction clause 139]
141. $false [avatar sat refutation 65,70,75,82,89,96,140] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:50,thDist:-50}
% SZS output end Proof for orders_1__t34_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 1
% 
% >>> Saturation
% Initial clauses: 22
% Generated clauses: 76
% Active clauses: 32
% Passive clauses: 32
% Final active clauses: 32
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Duplicate literals: 14
% Fw subsumption resolutions: 8
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 15
% Unit resulting resolution: 1
% Forward superposition: 11
% Backward superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward demodulation: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

