i: [1,99,0,2,1,1,0,0] t4_pre_topc
a: [3,166,0,2,1,99]
s: 166
i: [1,104,0,12,2,1,0,0] t4_pre_topc
a: [3,155,0,4,1,104]
s: 155
i: [1,100,0,5,1,1,0,0] t4_pre_topc
a: [3,161,0,5,1,100]
s: 161
a: [3,150,0,8,1,104]
s: 150
d: [2,188,0,7,1,2,22,150,155]
a: [3,192,0,7,1,188]
s: 192
i: [1,105,0,2,1,0,0,1] fc1_xboole_0
a: [3,171,0,2,1,105]
s: 171
i: [1,143,0,2,1,0,0,2] rc2_xboole_0
a: [3,176,0,2,1,143]
s: 176
i: [1,144,0,2,1,0,0,1] existence_l1_struct_0
a: [3,181,0,2,1,144]
s: 181
i: [1,145,0,2,1,0,0,2] rc1_xboole_0
a: [3,186,0,2,1,145]
s: 186
i: [1,106,0,3,1,0,0,1] fc1_subset_1
s: 106
i: [1,136,0,3,1,0,0,1] rc2_subset_1
s: 136
i: [1,139,0,3,1,0,0,3] reflexivity_r1_tarski
s: 139
i: [1,146,0,3,1,0,0,3] d10_xboole_0
s: 146
i: [1,147,0,3,1,0,0,3] d10_xboole_0
s: 147
i: [1,134,0,4,1,0,0,1] existence_m1_subset_1
s: 134
i: [1,138,0,4,1,0,0,2] rc3_subset_1
s: 138
i: [1,107,0,5,1,0,0,2] t4_boole
s: 107
i: [1,108,0,5,1,0,0,2] t3_boole
s: 108
i: [1,112,0,5,2,0,0,1] rc1_subset_1
s: 112
i: [1,118,0,5,2,0,0,1] t6_boole
s: 118
i: [1,131,0,5,2,0,0,255] t7_boole
s: 131
i: [1,135,0,5,1,0,0,1] rc2_subset_1
s: 135
i: [1,137,0,5,1,0,0,2] rc3_subset_1
s: 137
i: [1,114,0,6,2,0,0,2] rc4_subset_1
s: 114
i: [1,121,0,6,2,0,0,255] t1_subset
s: 121
i: [1,140,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 140
i: [1,111,0,7,2,0,0,1] rc1_subset_1
s: 111
i: [1,113,0,7,2,0,0,2] rc4_subset_1
s: 113
i: [1,115,0,7,2,0,0,1] fc12_struct_0
s: 115
i: [1,116,0,7,2,0,0,1] d3_struct_0
s: 116
i: [1,126,0,7,2,0,0,3] t3_subset
s: 126
d: [2,198,1,6,2,1,44,115,155]
d: [2,199,1,4,1,2,21,198,166]
a: [3,203,1,4,1,199]
s: 203
i: [1,127,0,7,2,0,0,3] t3_subset
s: 127
d: [2,205,1,6,2,1,44,116,155]
d: [2,209,1,4,1,2,21,205,166]
a: [3,213,1,4,1,209]
s: 213
r: 203
r: 192
r: 161
d: [2,231,0,4,1,2,23,161,213]
a: [3,246,0,4,1,231]
s: 246
d: [2,230,0,6,1,2,23,192,213]
a: [3,241,0,6,1,230]
s: 241
d: [2,229,1,3,1,2,23,203,213]
a: [3,236,1,3,1,229]
s: 236
i: [1,130,0,7,3,0,0,1] t8_boole
s: 130
d: [2,217,1,4,1,1,40,126,161]
a: [3,226,1,4,1,217]
s: 226
i: [1,117,0,8,2,0,0,1] dt_k2_struct_0
s: 117
i: [1,119,0,8,3,0,0,1] cc1_subset_1
s: 119
d: [2,221,1,4,1,0,40,126,135]
s: 221
i: [1,122,0,8,3,0,0,255] t2_subset
s: 122
r: 155
r: 213
r: 246
a: [3,154,0,4,1,104]
s: 154
a: [3,245,0,4,1,231]
s: 245
d: [2,293,1,0,0,1,52,139,245,127]
e: 293
d: [2,294,1,3,1,1,40,245,127]
d: [2,296,1,0,0,1,21,294,139]
e: 296
r: 245
s: 246
s: 203
s: 161
d: [2,303,1,5,2,1,40,161,119]
a: [3,309,1,3,1,303]
s: 309
s: 192
a: [3,151,0,8,1,104]
s: 151
a: [3,212,1,4,1,209]
s: 212
i: [1,128,0,8,2,0,0,2] t37_xboole_1
s: 128
s: 226
i: [1,129,0,8,2,0,0,2] t37_xboole_1
s: 129
i: [1,120,0,9,3,0,0,2] cc4_subset_1
s: 120
d: [2,272,1,7,2,1,40,122,246]
d: [2,280,1,4,1,1,21,272,106]
a: [3,284,1,4,1,280]
s: 284
i: [1,125,0,9,3,0,0,3] d10_xboole_0
s: 125
d: [2,344,2,4,1,1,40,284,140]
a: [3,350,2,4,1,344]
s: 350
i: [1,133,0,9,3,0,0,255] t5_subset
s: 133
i: [1,132,0,10,3,0,0,255] t4_subset
s: 132
d: [2,354,2,8,2,1,40,125,226]
d: [2,355,2,4,1,2,21,354,212]
a: [3,359,2,4,1,355]
s: 359
i: [1,109,0,11,4,0,0,2] cc3_subset_1
s: 109
a: [3,312,1,2,1,303]
s: 312
i: [1,110,0,11,4,0,0,2] cc2_subset_1
s: 110
i: [1,141,0,11,2,0,0,1] dt_k7_subset_1
s: 141
d: [2,305,1,9,2,1,40,161,122]
d: [2,316,1,5,1,1,21,305,106]
a: [3,320,1,5,1,316]
s: 320
i: [1,142,0,12,2,0,0,1] redefinition_k7_subset_1
s: 142
d: [2,423,1,9,2,1,44,241,142]
d: [2,433,1,5,1,2,21,423,246]
a: [3,437,1,5,1,433]
s: 437
d: [2,442,2,0,0,1,52,139,437,129]
e: 442
d: [2,443,2,6,2,1,44,437,129]
d: [2,444,2,3,1,1,19,443]
d: [2,446,2,0,0,1,21,444,139]
e: 446
r: 192
r: 241
r: 437
a: [3,240,0,6,1,230]
s: 240
a: [3,191,0,7,1,188]
s: 191
d: [2,450,1,8,2,1,44,141,240]
d: [2,455,1,4,1,2,21,450,246]
a: [3,459,1,4,1,455]
s: 459
r: 203
r: 312
s: 312
d: [2,470,2,7,3,1,40,459,109]
d: [2,478,2,5,2,2,21,470,171]
d: [2,479,2,3,1,3,21,478,312]
a: [3,483,2,3,1,479]
s: 483
a: [3,202,1,4,1,199]
s: 202
d: [2,403,1,11,2,1,44,141,151]
a: [3,412,1,5,1,403]
s: 412
r: 309
r: 312
r: 483
a: [3,313,1,2,1,303]
s: 313
s: 309
a: [3,482,2,3,1,479]
s: 482
a: [3,436,1,5,1,433]
s: 436
a: [3,408,1,6,1,403]
s: 408
d: [2,543,2,0,0,2,52,166,408,117]
e: 543
d: [2,544,2,2,1,1,40,408,117]
d: [2,548,2,0,0,2,21,544,166]
e: 548
r: 408
d: [2,475,2,3,1,1,40,459,126]
a: [3,488,2,3,1,475]
s: 488
a: [3,407,1,6,1,403]
s: 407
d: [2,472,2,5,2,1,40,459,133]
a: [3,501,2,3,1,472]
s: 501
d: [2,422,1,12,2,1,44,142,151]
a: [3,431,1,6,1,422]
s: 431
d: [2,373,1,7,2,1,40,132,161]
s: 373
d: [2,562,2,11,3,1,40,407,109]
d: [2,570,2,8,2,2,21,562,309]
a: [3,578,2,3,1,570]
s: 578
d: [2,322,1,10,2,1,44,151,116]
d: [2,323,1,8,1,2,21,322,166]
a: [3,327,1,8,1,323]
s: 327
d: [2,551,3,6,2,1,40,488,125]
a: [3,556,3,3,1,551]
s: 556
d: [2,329,1,6,2,0,44,128,107]
d: [2,331,1,3,1,0,19,329]
s: 331
r: 488
d: [2,222,1,4,1,0,40,126,137]
s: 222
d: [2,476,2,7,2,1,40,459,122]
d: [2,490,2,4,1,1,21,476,106]
a: [3,494,2,4,1,490]
s: 494
d: [2,220,1,5,1,0,40,126,134]
s: 220
d: [2,505,2,9,3,1,40,412,109]
d: [2,513,2,7,2,2,21,505,171]
d: [2,514,2,4,1,3,21,513,309]
a: [3,518,2,4,1,514]
s: 518
d: [2,269,1,5,2,0,40,119,137]
s: 269
d: [2,276,1,8,2,0,40,122,135]
d: [2,288,1,5,1,0,21,276,106]
s: 288
d: [2,510,2,4,1,1,40,412,126]
a: [3,523,2,4,1,510]
s: 523
d: [2,277,1,8,2,0,40,122,137]
d: [2,289,1,5,1,0,21,277,106]
s: 289
d: [2,599,2,7,2,1,44,128,431]
d: [2,600,2,4,1,1,19,599]
a: [3,610,2,4,1,600]
s: 610
d: [2,637,3,6,2,1,44,610,116]
d: [2,644,3,2,1,2,21,637,359]
d: [2,645,3,0,0,3,21,644,166]
e: 645
i: [1,101,0,12,2,1,0,0] t4_pre_topc
d: [2,290,1,6,2,1,44,154,116]
d: [2,291,1,4,1,2,21,290,166]
d: [2,304,1,4,1,1,40,161,126]
d: [2,424,1,10,2,1,44,192,142]
d: [2,439,1,5,1,2,21,424,161]
d: [2,425,1,12,2,1,44,151,142]
f: 645,600,548,446,425,439,355,304,291,188,99,100,101,104
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pre_topc__t4_pre_topc
% SZS output start Proof for pre_topc__t4_pre_topc
1. ! [X0] : (l1_struct_0(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (~(k2_struct_0(X0) = X1 & k7_subset_1(u1_struct_0(X0),k2_struct_0(X0),X1) != k1_xboole_0) & ~(k7_subset_1(u1_struct_0(X0),k2_struct_0(X0),X1) = k1_xboole_0 & k2_struct_0(X0) != X1)))) [input t4_pre_topc]
2. ~! [X0] : (l1_struct_0(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (~(k2_struct_0(X0) = X1 & k7_subset_1(u1_struct_0(X0),k2_struct_0(X0),X1) != k1_xboole_0) & ~(k7_subset_1(u1_struct_0(X0),k2_struct_0(X0),X1) = k1_xboole_0 & k2_struct_0(X0) != X1)))) [negated conjecture 1]
8. ! [X0] : ! [X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [input d10_xboole_0]
9. ! [X0] : (l1_struct_0(X0) => u1_struct_0(X0) = k2_struct_0(X0)) [input d3_struct_0]
12. ! [X0] : (l1_struct_0(X0) => m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0)))) [input dt_k2_struct_0]
29. ! [X0,X1,X2] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => k7_subset_1(X0,X1,X2) = k4_xboole_0(X1,X2)) [input redefinition_k7_subset_1]
30. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
33. ! [X0] : ! [X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [input t37_xboole_1]
35. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
44. ! [X0,X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [flattening 8]
45. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 35]
46. ! [X0,X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [flattening 33]
51. ! [X0] : r1_tarski(X0,X0) [rectify 30]
52. ? [X0] : (? [X1] : (((k2_struct_0(X0) = X1 & k7_subset_1(u1_struct_0(X0),k2_struct_0(X0),X1) != k1_xboole_0) | (k7_subset_1(u1_struct_0(X0),k2_struct_0(X0),X1) = k1_xboole_0 & k2_struct_0(X0) != X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_struct_0(X0)) [ennf transformation 2]
60. ! [X0] : (u1_struct_0(X0) = k2_struct_0(X0) | ~l1_struct_0(X0)) [ennf transformation 9]
61. ! [X0] : (m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0)) [ennf transformation 12]
75. ! [X0,X1,X2] : (k7_subset_1(X0,X1,X2) = k4_xboole_0(X1,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 29]
76. ? [X0] : (? [X1] : (((k2_struct_0(X0) = X1 & k7_subset_1(u1_struct_0(X0),k2_struct_0(X0),X1) != k1_xboole_0) | (k7_subset_1(u1_struct_0(X0),k2_struct_0(X0),X1) = k1_xboole_0 & k2_struct_0(X0) != X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_struct_0(X0)) => (? [X1] : (((k2_struct_0(sK0) = X1 & k1_xboole_0 != k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),X1)) | (k1_xboole_0 = k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),X1) & k2_struct_0(sK0) != X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_struct_0(sK0)) [choice axiom]
77. ? [X1] : (((k2_struct_0(sK0) = X1 & k1_xboole_0 != k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),X1)) | (k1_xboole_0 = k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),X1) & k2_struct_0(sK0) != X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => (((k2_struct_0(sK0) = sK1 & k1_xboole_0 != k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),sK1)) | (k1_xboole_0 = k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),sK1) & k2_struct_0(sK0) != sK1)) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
78. (((k2_struct_0(sK0) = sK1 & k1_xboole_0 != k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),sK1)) | (k1_xboole_0 = k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),sK1) & k2_struct_0(sK0) != sK1)) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_struct_0(sK0) [skolemisation 52,77,76]
83. ! [X0,X1] : ((X0 = X1 | (~r1_tarski(X1,X0) | ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [nnf transformation 44]
84. ! [X0,X1] : ((X0 = X1 | ~r1_tarski(X1,X0) | ~r1_tarski(X0,X1)) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [flattening 83]
85. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 45]
86. ! [X0,X1] : ((k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | k1_xboole_0 != k4_xboole_0(X0,X1))) [nnf transformation 46]
99. l1_struct_0(sK0) [cnf transformation 78] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 78] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
101. k1_xboole_0 != k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),sK1) | k2_struct_0(sK0) != sK1 [cnf transformation 78] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
104. k2_struct_0(sK0) = sK1 | k1_xboole_0 = k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),sK1) [cnf transformation 78] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
116. u1_struct_0(X0) = k2_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 60] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
117. m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) [cnf transformation 61] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
125. ~r1_tarski(X1,X0) | X0 = X1 | ~r1_tarski(X0,X1) [cnf transformation 84] {a:0,w:9,wCS:3600,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
126. ~m1_subset_1(X0,k1_zfmisc_1(X1)) | r1_tarski(X0,X1) [cnf transformation 85] {a:0,w:7,wCS:2800,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
128. k1_xboole_0 != k4_xboole_0(X0,X1) | r1_tarski(X0,X1) [cnf transformation 86] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
129. k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1) [cnf transformation 86] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
139. r1_tarski(X0,X0) [cnf transformation 51] {a:0,w:3,wCS:1200,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
142. k7_subset_1(X0,X1,X2) = k4_xboole_0(X1,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 75] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
149. 1 <=> k1_xboole_0 = k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),sK1) [avatar definition]
150. k1_xboole_0 != k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),sK1) <- (~1) [avatar component clause 149] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
151. k1_xboole_0 = k7_subset_1(u1_struct_0(sK0),k2_struct_0(sK0),sK1) <- (1) [avatar component clause 149] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
153. 2 <=> k2_struct_0(sK0) = sK1 [avatar definition]
154. k2_struct_0(sK0) != sK1 <- (~2) [avatar component clause 153] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
155. k2_struct_0(sK0) = sK1 <- (2) [avatar component clause 153] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
156. 1 | 2 [avatar split clause 104,153,149]
157. ~2 | ~1 [avatar split clause 101,149,153]
159. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
161. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (3) [avatar component clause 159] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
162. 3 [avatar split clause 100,159]
164. 4 <=> l1_struct_0(sK0) [avatar definition]
166. l1_struct_0(sK0) <- (4) [avatar component clause 164] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
167. 4 [avatar split clause 99,164]
188. k1_xboole_0 != k7_subset_1(u1_struct_0(sK0),sK1,sK1) <- (~1, 2) [forward demodulation 150,155] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
190. 9 <=> k1_xboole_0 = k7_subset_1(u1_struct_0(sK0),sK1,sK1) [avatar definition]
192. k1_xboole_0 != k7_subset_1(u1_struct_0(sK0),sK1,sK1) <- (~9) [avatar component clause 190] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
193. ~9 | 1 | ~2 [avatar split clause 188,153,149,190]
211. 11 <=> u1_struct_0(sK0) = sK1 [avatar definition]
212. u1_struct_0(sK0) != sK1 <- (~11) [avatar component clause 211] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
224. 12 <=> r1_tarski(sK1,u1_struct_0(sK0)) [avatar definition]
226. r1_tarski(sK1,u1_struct_0(sK0)) <- (12) [avatar component clause 224] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
290. u1_struct_0(sK0) != sK1 | ~l1_struct_0(sK0) <- (~2) [superposition 154,116] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
291. u1_struct_0(sK0) != sK1 <- (~2, 4) [subsumption resolution 290,166] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
292. ~11 | 2 | ~4 [avatar split clause 291,164,153,211]
304. r1_tarski(sK1,u1_struct_0(sK0)) <- (3) [resolution 161,126] {a:1,w:4,wCS:400,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
315. 12 | ~3 [avatar split clause 304,159,224]
354. u1_struct_0(sK0) = sK1 | ~r1_tarski(u1_struct_0(sK0),sK1) <- (12) [resolution 125,226] {a:2,w:8,wCS:800,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
355. ~r1_tarski(u1_struct_0(sK0),sK1) <- (~11, 12) [subsumption resolution 354,212] {a:2,w:4,wCS:400,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
357. 22 <=> r1_tarski(u1_struct_0(sK0),sK1) [avatar definition]
359. ~r1_tarski(u1_struct_0(sK0),sK1) <- (~22) [avatar component clause 357] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
360. ~22 | 11 | ~12 [avatar split clause 355,224,211,357]
406. 23 <=> m1_subset_1(k2_struct_0(sK0),k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
408. ~m1_subset_1(k2_struct_0(sK0),k1_zfmisc_1(u1_struct_0(sK0))) <- (~23) [avatar component clause 406] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
424. k1_xboole_0 != k4_xboole_0(sK1,sK1) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (~9) [superposition 192,142] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
425. k1_xboole_0 = k4_xboole_0(k2_struct_0(sK0),sK1) | ~m1_subset_1(k2_struct_0(sK0),k1_zfmisc_1(u1_struct_0(sK0))) <- (1) [superposition 151,142] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
429. 26 <=> k1_xboole_0 = k4_xboole_0(k2_struct_0(sK0),sK1) [avatar definition]
431. k1_xboole_0 = k4_xboole_0(k2_struct_0(sK0),sK1) <- (26) [avatar component clause 429] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
435. 27 <=> k1_xboole_0 = k4_xboole_0(sK1,sK1) [avatar definition]
437. k1_xboole_0 != k4_xboole_0(sK1,sK1) <- (~27) [avatar component clause 435] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
439. k1_xboole_0 != k4_xboole_0(sK1,sK1) <- (3, ~9) [subsumption resolution 424,161] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
440. ~27 | ~3 | 9 [avatar split clause 439,190,159,435]
441. ~23 | 26 | ~1 [avatar split clause 425,149,429,406]
443. k1_xboole_0 != k1_xboole_0 | ~r1_tarski(sK1,sK1) <- (~27) [superposition 437,129] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
444. ~r1_tarski(sK1,sK1) <- (~27) [trivial inequality removal 443] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
446. $false <- (~27) [subsumption resolution 444,139] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
447. 27 [avatar contradiction clause 446]
544. ~l1_struct_0(sK0) <- (~23) [resolution 408,117] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
548. $false <- (4, ~23) [subsumption resolution 544,166] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
549. ~4 | 23 [avatar contradiction clause 548]
599. k1_xboole_0 != k1_xboole_0 | r1_tarski(k2_struct_0(sK0),sK1) <- (26) [superposition 128,431] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
600. r1_tarski(k2_struct_0(sK0),sK1) <- (26) [trivial inequality removal 599] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
608. 45 <=> r1_tarski(k2_struct_0(sK0),sK1) [avatar definition]
610. r1_tarski(k2_struct_0(sK0),sK1) <- (45) [avatar component clause 608] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
611. 45 | ~26 [avatar split clause 600,429,608]
637. r1_tarski(u1_struct_0(sK0),sK1) | ~l1_struct_0(sK0) <- (45) [superposition 610,116] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
644. ~l1_struct_0(sK0) <- (~22, 45) [subsumption resolution 637,359] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
645. $false <- (4, ~22, 45) [subsumption resolution 644,166] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
646. ~4 | 22 | ~45 [avatar contradiction clause 645]
647. $false [avatar sat refutation 156,157,162,167,193,292,315,360,440,441,447,549,611,646] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:479,thDist:-479}
% SZS output end Proof for pre_topc__t4_pre_topc
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 40
% Introduced skolems: 10
% 
% >>> Saturation
% Initial clauses: 47
% Generated clauses: 366
% Active clauses: 93
% Passive clauses: 173
% Final active clauses: 76
% Final passive clauses: 71
% Main loop iterations started: 99
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Trivial inequalities: 4
% Fw subsumption resolutions: 52
% Fw demodulations: 5
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Equational tautologies: 1
% Forward subsumptions: 45
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 122
% Unit resulting resolution: 3
% Forward superposition: 31
% Backward superposition: 20
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 18
% Split components: 36
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 73
% SAT solver unit clauses: 10
% SAT solver binary clauses: 18
% 
% Memory used [KB]: 5245
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% property evaluation: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% forward subsumption resolution: 0.001 s
% superposition: 0.002 s
% other: 0.005 s ( own 0.0000 s ) 

