i: [1,44,0,3,1,1,0,0] t38_zfmisc_1
a: [3,73,0,3,1,44]
s: 73
i: [1,43,0,3,1,1,0,0] t38_zfmisc_1
a: [3,78,0,3,1,43]
s: 78
i: [1,42,0,3,1,1,0,0] t38_zfmisc_1
a: [3,83,0,3,1,42]
s: 83
i: [1,41,0,6,1,1,0,0] t38_zfmisc_1
a: [3,88,0,6,1,41]
s: 88
i: [1,45,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,93,0,2,1,45]
s: 93
i: [1,46,0,2,1,0,0,1] fc1_xboole_0
a: [3,98,0,2,1,46]
s: 98
i: [1,64,0,2,1,0,0,2] rc2_xboole_0
a: [3,108,0,2,1,64]
s: 108
i: [1,65,0,2,1,0,0,2] rc1_xboole_0
a: [3,113,0,2,1,65]
s: 113
i: [1,48,0,3,1,0,0,1] fc2_xboole_0
s: 48
i: [1,47,0,3,1,0,0,255] d2_xboole_0
a: [3,103,0,3,1,47]
s: 103
r: 93
i: [1,49,0,5,1,0,0,1] t1_boole
s: 49
i: [1,50,0,5,2,0,0,1] t6_boole
s: 50
i: [1,60,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 60
i: [1,62,0,6,2,0,0,1] fc4_xboole_0
s: 62
i: [1,63,0,6,2,0,0,1] fc5_xboole_0
s: 63
i: [1,51,0,7,3,0,0,1] t8_boole
s: 51
i: [1,61,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 61
i: [1,52,0,17,4,0,0,1] t37_zfmisc_1
d: [2,66,0,13,3,0,18,52]
s: 66
d: [2,133,1,12,3,1,44,66,88]
d: [2,138,1,9,2,2,21,133,78]
s: 138
i: [1,59,0,17,4,0,0,1] t37_zfmisc_1
d: [2,69,0,13,3,0,18,59]
s: 69
i: [1,54,0,18,4,0,0,1] t37_zfmisc_1
d: [2,67,0,14,3,0,18,54]
s: 67
d: [2,139,2,4,1,2,47,138]
a: [3,143,2,4,1,139]
s: 143
r: 138
r: 88
d: [2,158,0,5,1,2,23,88,143]
a: [3,163,0,5,1,158]
s: 163
d: [2,159,3,2,1,1,44,48,143]
a: [3,168,3,2,1,159]
s: 168
i: [1,55,0,18,4,0,0,1] t37_zfmisc_1
d: [2,68,0,14,3,0,18,55]
s: 68
i: [1,56,0,16,4,0,0,1] t37_zfmisc_1
s: 56
d: [2,174,1,11,3,1,44,69,163]
d: [2,180,1,8,2,2,21,174,73]
s: 180
d: [2,181,3,6,2,3,44,180,143]
d: [2,182,3,3,1,3,19,181]
d: [2,183,3,0,0,4,21,182,83]
e: 183
f: 183,158,139,41,42,43,44
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t38_zfmisc_1
% SZS output start Proof for zfmisc_1__t38_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ~(k1_xboole_0 != X2 & k1_xboole_0 != X1 & X1 != X2 & k1_tarski(X0) = k2_xboole_0(X1,X2)) [input t38_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ~(k1_xboole_0 != X2 & k1_xboole_0 != X1 & X1 != X2 & k1_tarski(X0) = k2_xboole_0(X1,X2)) [negated conjecture 1]
18. ! [X0] : ! [X1] : ! [X2] : ~(~(k1_xboole_0 = X2 & k1_tarski(X0) = X1) & ~(k1_tarski(X0) = X2 & k1_xboole_0 = X1) & ~(k1_tarski(X0) = X2 & k1_tarski(X0) = X1) & k1_tarski(X0) = k2_xboole_0(X1,X2)) [input t37_zfmisc_1]
22. ~! [X0,X1,X2] : ~(k1_xboole_0 != X2 & k1_xboole_0 != X1 & X1 != X2 & k1_tarski(X0) = k2_xboole_0(X1,X2)) [flattening 2]
25. ! [X0,X1,X2] : ~(~(k1_xboole_0 = X2 & k1_tarski(X0) = X1) & ~(k1_tarski(X0) = X2 & k1_xboole_0 = X1) & ~(k1_tarski(X0) = X2 & k1_tarski(X0) = X1) & k1_tarski(X0) = k2_xboole_0(X1,X2)) [flattening 18]
29. ? [X0,X1,X2] : (k1_xboole_0 != X2 & k1_xboole_0 != X1 & X1 != X2 & k1_tarski(X0) = k2_xboole_0(X1,X2)) [ennf transformation 22]
32. ! [X0,X1,X2] : ((k1_xboole_0 = X2 & k1_tarski(X0) = X1) | (k1_tarski(X0) = X2 & k1_xboole_0 = X1) | (k1_tarski(X0) = X2 & k1_tarski(X0) = X1) | k1_tarski(X0) != k2_xboole_0(X1,X2)) [ennf transformation 25]
35. ? [X0,X1,X2] : (k1_xboole_0 != X2 & k1_xboole_0 != X1 & X1 != X2 & k1_tarski(X0) = k2_xboole_0(X1,X2)) => (k1_xboole_0 != sK2 & k1_xboole_0 != sK1 & sK1 != sK2 & k1_tarski(sK0) = k2_xboole_0(sK1,sK2)) [choice axiom]
36. k1_xboole_0 != sK2 & k1_xboole_0 != sK1 & sK1 != sK2 & k1_tarski(sK0) = k2_xboole_0(sK1,sK2) [skolemisation 29,35]
41. k1_tarski(sK0) = k2_xboole_0(sK1,sK2) [cnf transformation 36] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. sK1 != sK2 [cnf transformation 36] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. k1_xboole_0 != sK1 [cnf transformation 36] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
44. k1_xboole_0 != sK2 [cnf transformation 36] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. k1_tarski(X0) = X1 | k1_xboole_0 = X1 | k1_tarski(X0) = X1 | k1_tarski(X0) != k2_xboole_0(X1,X2) [cnf transformation 32] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
59. k1_xboole_0 = X2 | k1_tarski(X0) = X2 | k1_tarski(X0) = X2 | k1_tarski(X0) != k2_xboole_0(X1,X2) [cnf transformation 32] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
66. k1_tarski(X0) != k2_xboole_0(X1,X2) | k1_xboole_0 = X1 | k1_tarski(X0) = X1 [duplicate literal removal 52] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. k1_tarski(X0) != k2_xboole_0(X1,X2) | k1_tarski(X0) = X2 | k1_xboole_0 = X2 [duplicate literal removal 59] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. 1 <=> k1_xboole_0 = sK2 [avatar definition]
73. k1_xboole_0 != sK2 <- (~1) [avatar component clause 71] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
74. ~1 [avatar split clause 44,71]
76. 2 <=> k1_xboole_0 = sK1 [avatar definition]
78. k1_xboole_0 != sK1 <- (~2) [avatar component clause 76] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. ~2 [avatar split clause 43,76]
81. 3 <=> sK1 = sK2 [avatar definition]
83. sK1 != sK2 <- (~3) [avatar component clause 81] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. ~3 [avatar split clause 42,81]
86. 4 <=> k1_tarski(sK0) = k2_xboole_0(sK1,sK2) [avatar definition]
88. k1_tarski(sK0) = k2_xboole_0(sK1,sK2) <- (4) [avatar component clause 86] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
89. 4 [avatar split clause 41,86]
133. k1_tarski(X0) != k1_tarski(sK0) | k1_xboole_0 = sK1 | k1_tarski(X0) = sK1 <- (4) [superposition 66,88] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
138. k1_tarski(X0) != k1_tarski(sK0) | k1_tarski(X0) = sK1 <- (~2, 4) [subsumption resolution 133,78] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
139. sK1 = k1_tarski(sK0) <- (~2, 4) [equality resolution 138] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
141. 10 <=> sK1 = k1_tarski(sK0) [avatar definition]
143. sK1 = k1_tarski(sK0) <- (10) [avatar component clause 141] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. 10 | 2 | ~4 [avatar split clause 139,86,76,141]
158. sK1 = k2_xboole_0(sK1,sK2) <- (4, 10) [backward demodulation 88,143] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
161. 11 <=> sK1 = k2_xboole_0(sK1,sK2) [avatar definition]
163. sK1 = k2_xboole_0(sK1,sK2) <- (11) [avatar component clause 161] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
164. 11 | ~4 | ~10 [avatar split clause 158,141,86,161]
174. sK1 != k1_tarski(X2) | sK2 = k1_tarski(X2) | k1_xboole_0 = sK2 <- (11) [superposition 69,163] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
180. sK1 != k1_tarski(X2) | sK2 = k1_tarski(X2) <- (~1, 11) [subsumption resolution 174,73] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
181. sK1 != sK1 | sK1 = sK2 <- (~1, 10, 11) [superposition 180,143] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
182. sK1 = sK2 <- (~1, 10, 11) [trivial inequality removal 181] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
183. $false <- (~1, ~3, 10, 11) [subsumption resolution 182,83] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
184. 1 | 3 | ~10 | ~11 [avatar contradiction clause 183]
185. $false [avatar sat refutation 74,79,84,89,144,164,184] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:53,thDist:-53}
% SZS output end Proof for zfmisc_1__t38_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 25
% Generated clauses: 89
% Active clauses: 25
% Passive clauses: 39
% Final active clauses: 22
% Final passive clauses: 12
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Trivial inequalities: 1
% Fw subsumption resolutions: 4
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 15
% 
% >>> Generating Inferences
% Forward superposition: 23
% Backward superposition: 14
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 9
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

