i: [1,224,0,2,1,1,0,0] t135_xreal_1
a: [3,353,0,2,1,224]
s: 353
i: [1,223,0,2,1,1,0,0] t135_xreal_1
a: [3,358,0,2,1,223]
s: 358
i: [1,226,0,3,1,1,0,0] t135_xreal_1
a: [3,343,0,3,1,226]
s: 343
i: [1,225,0,3,1,1,0,0] t135_xreal_1
a: [3,348,0,3,1,225]
s: 348
i: [1,227,0,5,1,1,0,0] t135_xreal_1
a: [3,338,0,5,1,227]
s: 338
i: [1,228,0,2,1,0,0,255] fc1_numbers
a: [3,363,0,2,1,228]
s: 363
i: [1,229,0,2,1,0,0,2] fc1_xboole_0
a: [3,368,0,2,1,229]
s: 368
i: [1,236,0,2,1,0,0,3] fc6_ordinal1
a: [3,398,0,2,1,236]
s: 398
i: [1,235,0,2,1,0,0,3] fc6_ordinal1
a: [3,403,0,2,1,235]
s: 403
i: [1,302,0,2,1,0,0,2] rc2_xboole_0
a: [3,415,0,2,1,302]
s: 415
i: [1,303,0,2,1,0,0,4] rc4_ordinal1
a: [3,420,0,2,1,303]
s: 420
i: [1,304,0,2,1,0,0,4] rc1_ordinal1
a: [3,425,0,2,1,304]
s: 425
i: [1,305,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,430,0,2,1,305]
s: 430
i: [1,306,0,2,1,0,0,1] rc1_xreal_0
a: [3,435,0,2,1,306]
s: 435
i: [1,307,0,2,1,0,0,2] rc1_xxreal_0
a: [3,440,0,2,1,307]
s: 440
i: [1,308,0,2,1,0,0,2] rc1_xboole_0
a: [3,445,0,2,1,308]
s: 445
i: [1,312,0,2,1,0,0,255] rc2_ordinal1
a: [3,450,0,2,1,312]
s: 450
i: [1,311,0,2,1,0,0,255] rc2_ordinal1
a: [3,455,0,2,1,311]
s: 455
i: [1,310,0,2,1,0,0,255] rc2_ordinal1
a: [3,460,0,2,1,310]
s: 460
i: [1,309,0,2,1,0,0,255] rc2_ordinal1
a: [3,465,0,2,1,309]
s: 465
i: [1,314,0,2,1,0,0,4] rc5_ordinal1
a: [3,470,0,2,1,314]
s: 470
i: [1,313,0,2,1,0,0,4] rc5_ordinal1
a: [3,475,0,2,1,313]
s: 475
i: [1,316,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,480,0,2,1,316]
s: 480
i: [1,315,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,485,0,2,1,315]
s: 485
i: [1,320,0,2,1,0,0,2] rc2_xreal_0
a: [3,490,0,2,1,320]
s: 490
i: [1,319,0,2,1,0,0,2] rc2_xreal_0
a: [3,495,0,2,1,319]
s: 495
i: [1,318,0,2,1,0,0,2] rc2_xreal_0
a: [3,500,0,2,1,318]
s: 500
i: [1,317,0,2,1,0,0,2] rc2_xreal_0
a: [3,505,0,2,1,317]
s: 505
i: [1,324,0,2,1,0,0,2] rc3_xreal_0
a: [3,510,0,2,1,324]
s: 510
i: [1,323,0,2,1,0,0,2] rc3_xreal_0
a: [3,515,0,2,1,323]
s: 515
i: [1,322,0,2,1,0,0,2] rc3_xreal_0
a: [3,520,0,2,1,322]
s: 520
i: [1,321,0,2,1,0,0,2] rc3_xreal_0
a: [3,525,0,2,1,321]
s: 525
i: [1,326,0,2,1,0,0,2] rc2_xxreal_0
a: [3,530,0,2,1,326]
s: 530
i: [1,325,0,2,1,0,0,2] rc2_xxreal_0
a: [3,535,0,2,1,325]
s: 535
i: [1,328,0,2,1,0,0,2] rc3_xxreal_0
a: [3,540,0,2,1,328]
s: 540
i: [1,327,0,2,1,0,0,2] rc3_xxreal_0
a: [3,545,0,2,1,327]
s: 545
i: [1,332,0,2,1,0,0,2] rc4_xreal_0
a: [3,550,0,2,1,332]
s: 550
i: [1,331,0,2,1,0,0,2] rc4_xreal_0
a: [3,555,0,2,1,331]
s: 555
i: [1,330,0,2,1,0,0,2] rc4_xreal_0
a: [3,560,0,2,1,330]
s: 560
i: [1,329,0,2,1,0,0,2] rc4_xreal_0
a: [3,565,0,2,1,329]
s: 565
i: [1,334,0,2,1,0,0,2] rc4_xxreal_0
a: [3,570,0,2,1,334]
s: 570
i: [1,333,0,2,1,0,0,2] rc4_xxreal_0
a: [3,575,0,2,1,333]
s: 575
i: [1,287,0,3,1,0,0,255] reflexivity_r1_tarski
s: 287
i: [1,230,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,373,0,3,1,230]
s: 373
i: [1,231,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,378,0,3,1,231]
s: 378
r: 368
d: [2,577,0,2,1,2,23,368,378]
a: [3,581,0,2,1,577]
s: 581
i: [1,232,0,3,1,0,0,2] t1_numerals
a: [3,383,0,3,1,232]
s: 383
d: [2,583,0,3,1,2,22,383,378]
a: [3,587,0,3,1,583]
s: 587
i: [1,238,0,4,2,0,0,4] cc6_ordinal1
s: 238
i: [1,239,0,4,2,0,0,4] cc2_xcmplx_0
s: 239
i: [1,240,0,4,2,0,0,4] cc2_xreal_0
s: 240
d: [2,589,1,2,1,1,40,238,420]
a: [3,594,1,2,1,589]
s: 594
i: [1,241,0,4,2,0,0,4] cc2_xxreal_0
s: 241
d: [2,590,1,2,1,1,40,238,470]
a: [3,599,1,2,1,590]
s: 599
i: [1,242,0,4,2,0,0,255] cc1_ordinal1
s: 242
i: [1,243,0,4,2,0,0,255] cc1_ordinal1
s: 243
d: [2,601,1,2,1,1,40,239,420]
a: [3,606,1,2,1,601]
s: 606
i: [1,246,0,4,2,0,0,2] cc3_xreal_0
s: 246
d: [2,687,1,2,1,1,40,246,358]
a: [3,696,1,2,1,687]
s: 696
i: [1,247,0,4,2,0,0,2] cc4_xreal_0
s: 247
i: [1,256,0,4,2,0,0,255] cc9_ordinal1
s: 256
d: [2,688,1,2,1,1,40,246,353]
a: [3,701,1,2,1,688]
s: 701
i: [1,257,0,4,2,0,0,4] cc7_ordinal1
s: 257
d: [2,708,1,2,1,1,40,247,358]
a: [3,717,1,2,1,708]
s: 717
i: [1,258,0,4,2,0,0,4] cc3_ordinal1
s: 258
i: [1,286,0,4,1,0,0,3] existence_m1_subset_1
s: 286
d: [2,709,1,2,1,1,40,247,353]
a: [3,722,1,2,1,709]
s: 722
i: [1,233,0,4,1,0,0,255] dt_k5_numbers
a: [3,388,0,4,1,233]
s: 388
d: [2,602,1,2,1,1,40,239,470]
a: [3,611,1,2,1,602]
s: 611
i: [1,234,0,4,1,0,0,255] dt_k6_numbers
a: [3,393,0,4,1,234]
s: 393
d: [2,801,0,4,1,2,22,388,373]
a: [3,805,0,4,1,801]
s: 805
d: [2,613,1,2,1,1,40,240,420]
a: [3,618,1,2,1,613]
s: 618
d: [2,807,0,4,1,2,22,393,373]
a: [3,811,0,4,1,807]
s: 811
d: [2,614,1,2,1,1,40,240,470]
a: [3,623,1,2,1,614]
s: 623
i: [1,259,0,5,2,0,0,2] t6_boole
s: 259
i: [1,260,0,5,2,0,0,3] cc8_ordinal1
s: 260
d: [2,625,1,2,1,1,40,241,420]
a: [3,630,1,2,1,625]
s: 630
i: [1,261,0,5,2,0,0,255] cc1_xcmplx_0
s: 261
d: [2,626,1,2,1,1,40,241,470]
a: [3,635,1,2,1,626]
s: 635
i: [1,262,0,5,2,0,0,255] cc1_xreal_0
s: 262
i: [1,283,0,5,2,0,0,255] t7_boole
s: 283
d: [2,637,1,2,1,1,40,242,398]
a: [3,645,1,2,1,637]
s: 645
i: [1,300,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,410,0,5,2,300]
s: 410
d: [2,639,1,2,1,1,40,242,425]
a: [3,655,1,2,1,639]
s: 655
d: [2,819,0,5,2,1,22,259,378]
s: 819
i: [1,267,0,6,3,0,0,255] cc2_ordinal1
s: 267
d: [2,842,1,5,2,2,44,343,819]
s: 842
i: [1,270,0,6,3,0,0,2] cc5_xxreal_0
s: 270
d: [2,843,1,5,2,2,44,348,819]
s: 843
i: [1,272,0,6,3,0,0,2] cc5_xxreal_0
s: 272
i: [1,273,0,6,3,0,0,2] cc3_xxreal_0
s: 273
d: [2,841,1,7,2,2,44,338,819]
s: 841
i: [1,275,0,6,3,0,0,2] cc3_xxreal_0
s: 275
d: [2,890,2,8,2,3,40,841,410]
a: [3,900,2,4,1,890]
s: 900
i: [1,277,0,6,3,0,0,2] cc7_xxreal_0
s: 277
i: [1,278,0,6,3,0,0,2] cc7_xxreal_0
s: 278
d: [2,662,1,2,1,1,40,243,398]
a: [3,670,1,2,1,662]
s: 670
i: [1,279,0,6,2,0,0,255] t1_subset
s: 279
d: [2,664,1,2,1,1,40,243,425]
a: [3,680,1,2,1,664]
s: 680
i: [1,288,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 288
i: [1,244,0,7,3,0,0,4] cc5_ordinal1
s: 244
d: [2,689,1,2,1,1,40,246,435]
a: [3,706,1,2,1,689]
s: 706
i: [1,245,0,7,2,0,0,1] t5_arithm
s: 245
d: [2,710,1,2,1,1,40,247,435]
a: [3,727,1,2,1,710]
s: 727
i: [1,281,0,7,2,0,0,255] t3_subset
s: 281
i: [1,282,0,7,3,0,0,2] t8_boole
s: 282
d: [2,729,1,2,1,1,40,256,581]
a: [3,736,1,2,1,729]
s: 736
i: [1,237,0,8,3,0,0,255] cc10_ordinal1
s: 237
d: [2,730,1,2,1,1,40,256,445]
a: [3,741,1,2,1,730]
s: 741
i: [1,264,0,8,4,0,0,2] cc4_xxreal_0
s: 264
i: [1,266,0,8,4,0,0,2] cc6_xxreal_0
s: 266
d: [2,891,2,6,2,3,40,841,842]
a: [3,904,2,2,1,891]
s: 904
r: 904
i: [1,268,0,8,4,0,0,2] cc8_xxreal_0
s: 268
a: [3,905,2,2,1,891]
s: 905
i: [1,280,0,8,3,0,0,255] t2_subset
s: 280
i: [1,298,0,8,3,0,0,1] fc5_xcmplx_0
s: 298
d: [2,892,2,6,2,3,40,841,843]
a: [3,909,2,2,1,892]
s: 909
r: 909
i: [1,299,0,8,3,0,0,1] fc8_xreal_0
s: 299
a: [3,910,2,2,1,892]
s: 910
i: [1,285,0,9,3,0,0,255] t5_subset
s: 285
i: [1,284,0,10,3,0,0,255] t4_subset
s: 284
d: [2,934,2,6,3,1,40,264,717]
a: [3,949,2,2,1,934]
s: 949
i: [1,301,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 301
d: [2,935,2,6,3,1,40,264,722]
a: [3,958,2,2,1,935]
s: 958
i: [1,250,0,11,5,0,0,1] t2_real
s: 250
i: [1,251,0,11,5,0,0,1] t1_real
s: 251
d: [2,1098,1,8,4,1,40,250,343]
d: [2,1105,1,6,3,2,21,1098,353]
a: [3,1109,1,2,1,1105]
s: 1109
i: [1,252,0,11,5,0,0,1] t8_real
s: 252
a: [3,953,2,2,1,934]
s: 953
d: [2,1168,3,0,0,3,52,717,949,953,272]
e: 1168
d: [2,1169,3,4,2,1,40,953,272]
d: [2,1172,3,2,1,2,21,1169,949]
d: [2,1173,3,0,0,3,21,1172,717]
e: 1173
r: 953
i: [1,253,0,11,5,0,0,1] t7_real
s: 253
i: [1,254,0,11,5,0,0,1] t4_real
s: 254
a: [3,962,2,2,1,935]
s: 962
d: [2,1175,3,0,0,3,52,722,958,962,272]
e: 1175
d: [2,1176,3,4,2,1,40,962,272]
d: [2,1179,3,2,1,2,21,1176,958]
d: [2,1180,3,0,0,3,21,1179,722]
e: 1180
r: 962
i: [1,255,0,11,5,0,0,1] t3_real
s: 255
a: [3,896,2,4,1,890]
s: 896
i: [1,289,0,12,5,0,0,1] fc31_xreal_0
s: 289
i: [1,290,0,12,5,0,0,1] fc32_xreal_0
s: 290
d: [2,1149,2,4,2,2,44,1109,819]
s: 1149
d: [2,1184,3,0,0,4,52,550,565,1149]
e: 1184
d: [2,1187,3,2,1,3,40,1149,565]
d: [2,1190,3,0,0,4,21,1187,550]
e: 1190
r: 949
r: 958
r: 1149
r: 1109
i: [1,291,0,12,5,0,0,1] fc34_xreal_0
s: 291
a: [3,1108,1,2,1,1105]
s: 1108
i: [1,292,0,12,5,0,0,1] fc33_xreal_0
s: 292
i: [1,297,0,12,5,0,0,1] fc9_xcmplx_0
s: 297
d: [2,1080,1,11,3,1,40,301,338]
a: [3,1086,1,2,1,1080]
s: 1086
i: [1,248,0,13,6,0,0,1] t6_real
s: 248
a: [3,1113,1,2,1,1105]
s: 1113
i: [1,249,0,13,6,0,0,1] t5_real
s: 249
i: [1,296,0,14,4,0,0,255] existence_m2_subset_1
s: 296
d: [2,1126,1,8,4,1,40,251,343]
d: [2,1133,1,6,3,2,21,1126,958]
d: [2,1134,1,4,2,3,21,1133,353]
a: [3,1137,1,2,1,1134]
s: 1137
i: [1,293,0,15,5,0,0,255] dt_m2_subset_1
s: 293
a: [3,948,2,2,1,934]
s: 948
i: [1,294,0,15,5,0,0,255] redefinition_m2_subset_1
s: 294
i: [1,295,0,15,5,0,0,255] redefinition_m2_subset_1
s: 295
s: 953
d: [2,731,1,2,1,1,40,256,565]
a: [3,746,1,2,1,731]
s: 746
a: [3,957,2,2,1,935]
s: 957
d: [2,732,1,2,1,1,40,256,575]
a: [3,751,1,2,1,732]
s: 751
d: [2,753,1,2,1,1,40,257,581]
a: [3,760,1,2,1,753]
s: 760
s: 962
d: [2,754,1,2,1,1,40,257,445]
a: [3,765,1,2,1,754]
s: 765
d: [2,1210,2,2,1,1,40,1108,246]
a: [3,1216,2,2,1,1210]
s: 1216
d: [2,755,1,2,1,1,40,257,565]
a: [3,770,1,2,1,755]
s: 770
d: [2,756,1,2,1,1,40,257,575]
a: [3,775,1,2,1,756]
s: 775
d: [2,1151,2,14,5,2,40,252,841]
a: [3,1159,2,4,1,1151]
s: 1159
d: [2,1336,3,0,0,3,52,358,353,1159,299]
e: 1336
d: [2,1337,3,4,2,1,40,1159,299]
d: [2,1339,3,2,1,2,21,1337,353]
d: [2,1340,3,0,0,3,21,1339,358]
e: 1340
r: 1113
r: 1137
r: 1159
a: [3,1112,1,2,1,1105]
s: 1112
d: [2,1342,2,0,0,3,52,1086,581,1112,270]
e: 1342
d: [2,1344,2,4,2,1,40,1112,270]
d: [2,1348,2,2,1,2,21,1344,581]
d: [2,1349,2,0,0,3,21,1348,1086]
e: 1349
r: 1112
s: 1113
s: 1137
d: [2,777,1,2,1,1,40,258,581]
a: [3,784,1,2,1,777]
s: 784
a: [3,1158,2,4,1,1151]
s: 1158
d: [2,1366,3,0,0,2,52,896,1158,247]
e: 1366
d: [2,1367,3,4,1,1,40,1158,247]
d: [2,1370,3,0,0,2,21,1367,896]
e: 1370
r: 896
a: [3,1091,1,5,1,1080]
s: 1091
a: [3,895,2,4,1,890]
s: 895
d: [2,778,1,2,1,1,40,258,445]
a: [3,789,1,2,1,778]
s: 789
d: [2,779,1,2,1,1,40,258,565]
a: [3,794,1,2,1,779]
s: 794
a: [3,1166,2,4,1,1151]
s: 1166
d: [2,1423,3,0,0,5,52,948,358,957,353,1166,291]
e: 1423
d: [2,1425,3,8,4,1,40,1166,291]
d: [2,1429,3,6,3,2,21,1425,353]
d: [2,1430,3,4,2,3,21,1429,957]
d: [2,1431,3,2,1,4,21,1430,358]
d: [2,1432,3,0,0,5,21,1431,948]
e: 1432
r: 948
r: 953
r: 1137
a: [3,1138,1,2,1,1134]
s: 1138
d: [2,1440,2,0,0,3,52,1086,581,1138,273]
e: 1440
d: [2,1441,2,4,2,1,40,1138,273]
d: [2,1444,2,2,1,2,21,1441,581]
d: [2,1445,2,0,0,3,21,1444,1086]
e: 1445
d: [2,1127,1,8,4,1,40,251,348]
d: [2,1194,1,6,3,2,21,1127,358]
d: [2,1198,1,6,3,2,21,1126,353]
d: [2,1150,1,12,4,1,40,252,338]
d: [2,1209,2,2,1,1,40,1108,247]
f: 1445,1431,1340,1209,1150,1198,1194,1190,577,329,332,231,229,223,224,225,226,227
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t135_xreal_1
% SZS output start Proof for xreal_1__t135_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,k6_numbers)) => r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1))))) [input t135_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,k6_numbers)) => r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1))))) [negated conjecture 1]
12. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
22. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc34_xreal_0]
30. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
40. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
58. ? [X0] : (v1_xreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0) & v1_xboole_0(X0)) [input rc4_xreal_0]
73. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
79. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
82. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
86. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
87. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
100. ? [X0] : (? [X1] : ((~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1)) & (r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,k6_numbers))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
101. ? [X0] : (? [X1] : (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1)) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 100]
111. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 86]
118. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 30]
119. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 118]
120. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 40]
121. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 120]
131. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 79]
145. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 12]
146. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 145]
163. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 22]
164. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 163]
177. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 87]
178. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 177]
183. ? [X0] : (? [X1] : (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1)) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,X1)) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
184. ? [X1] : (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,X1)) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) => (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) & r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) [choice axiom]
185. (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) & r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 101,184,183]
219. ? [X0] : (v1_xreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0) & v1_xboole_0(X0)) => (v1_xreal_0(sK18) & v1_xxreal_0(sK18) & v1_xcmplx_0(sK18) & v1_xboole_0(sK18)) [choice axiom]
220. v1_xreal_0(sK18) & v1_xxreal_0(sK18) & v1_xcmplx_0(sK18) & v1_xboole_0(sK18) [skolemisation 58,219]
223. v1_xreal_0(sK0) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
224. v1_xreal_0(sK1) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
225. r1_xxreal_0(sK0,k6_numbers) [cnf transformation 185] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
226. r1_xxreal_0(sK1,k6_numbers) [cnf transformation 185] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
227. ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) [cnf transformation 185] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
229. v1_xboole_0(k1_xboole_0) [cnf transformation 73] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
231. k6_numbers = k1_xboole_0 [cnf transformation 82] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
247. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 111] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
251. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 119] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
252. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 121] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
259. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 131] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
273. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 146] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
291. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 164] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
299. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 178] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
329. v1_xboole_0(sK18) [cnf transformation 220] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
332. v1_xreal_0(sK18) [cnf transformation 220] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
336. 1 <=> r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) [avatar definition]
338. ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) <- (~1) [avatar component clause 336] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
339. ~1 [avatar split clause 227,336]
341. 2 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
343. r1_xxreal_0(sK1,k6_numbers) <- (2) [avatar component clause 341] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
344. 2 [avatar split clause 226,341]
346. 3 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
348. r1_xxreal_0(sK0,k6_numbers) <- (3) [avatar component clause 346] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
349. 3 [avatar split clause 225,346]
351. 4 <=> v1_xreal_0(sK1) [avatar definition]
353. v1_xreal_0(sK1) <- (4) [avatar component clause 351] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
354. 4 [avatar split clause 224,351]
356. 5 <=> v1_xreal_0(sK0) [avatar definition]
358. v1_xreal_0(sK0) <- (5) [avatar component clause 356] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
359. 5 [avatar split clause 223,356]
366. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
368. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 366] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
369. 7 [avatar split clause 229,366]
376. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
378. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 376] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
379. 9 [avatar split clause 231,376]
548. 44 <=> v1_xreal_0(sK18) [avatar definition]
550. v1_xreal_0(sK18) <- (44) [avatar component clause 548] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
551. 44 [avatar split clause 332,548]
563. 47 <=> v1_xboole_0(sK18) [avatar definition]
565. v1_xboole_0(sK18) <- (47) [avatar component clause 563] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
566. 47 [avatar split clause 329,563]
577. v1_xboole_0(k6_numbers) <- (7, 9) [backward demodulation 368,378] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
579. 50 <=> v1_xboole_0(k6_numbers) [avatar definition]
581. v1_xboole_0(k6_numbers) <- (50) [avatar component clause 579] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
582. 50 | ~7 | ~9 [avatar split clause 577,376,366,579]
819. k6_numbers = X0 | ~v1_xboole_0(X0) <- (9) [forward demodulation 259,378] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
947. 103 <=> v2_xxreal_0(sK0) [avatar definition]
956. 105 <=> v2_xxreal_0(sK1) [avatar definition]
957. ~v2_xxreal_0(sK1) <- (~105) [avatar component clause 956] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1085. 127 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1086. v1_xxreal_0(k6_numbers) <- (127) [avatar component clause 1085] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1107. 130 <=> v1_xreal_0(k6_numbers) [avatar definition]
1108. v1_xreal_0(k6_numbers) <- (130) [avatar component clause 1107] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1109. ~v1_xreal_0(k6_numbers) <- (~130) [avatar component clause 1107] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1126. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (2) [resolution 251,343] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1127. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (3) [resolution 251,348] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1136. 133 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1138. v2_xxreal_0(k6_numbers) <- (133) [avatar component clause 1136] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1149. ~v1_xboole_0(X0) | ~v1_xreal_0(X0) <- (9, ~130) [superposition 1109,819] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1150. v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (~1) [resolution 252,338] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1157. 134 <=> v1_xreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1159. ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~134) [avatar component clause 1157] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1164. 136 <=> v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1166. v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) <- (136) [avatar component clause 1164] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1187. ~v1_xreal_0(sK18) <- (9, 47, ~130) [resolution 1149,565] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
1190. $false <- (9, 44, 47, ~130) [subsumption resolution 1187,550] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
1191. ~9 | ~44 | ~47 | 130 [avatar contradiction clause 1190]
1194. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 5) [subsumption resolution 1127,358] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1195. ~130 | 133 | ~103 | ~3 | ~5 [avatar split clause 1194,356,346,947,1136,1107]
1198. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 4) [subsumption resolution 1126,353] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1199. ~130 | 133 | ~105 | ~2 | ~4 [avatar split clause 1198,351,341,956,1136,1107]
1205. ~130 | ~134 | 133 | 136 | 1 [avatar split clause 1150,336,1164,1136,1157,1107]
1209. v1_xxreal_0(k6_numbers) <- (130) [resolution 1108,247] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1212. 127 | ~130 [avatar split clause 1209,1107,1085]
1337. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~134) [resolution 1159,299] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1339. ~v1_xreal_0(sK0) <- (4, ~134) [subsumption resolution 1337,353] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1340. $false <- (4, 5, ~134) [subsumption resolution 1339,358] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1341. ~4 | ~5 | 134 [avatar contradiction clause 1340]
1425. ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (136) [resolution 1166,291] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1429. v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (4, 136) [subsumption resolution 1425,353] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1430. ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (4, ~105, 136) [subsumption resolution 1429,957] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1431. v2_xxreal_0(sK0) <- (4, 5, ~105, 136) [subsumption resolution 1430,358] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1435. 103 | ~4 | ~5 | 105 | ~136 [avatar split clause 1431,1164,956,356,351,947]
1441. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (133) [resolution 1138,273] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1444. ~v1_xxreal_0(k6_numbers) <- (50, 133) [subsumption resolution 1441,581] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1445. $false <- (50, 127, 133) [subsumption resolution 1444,1086] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1446. ~50 | ~127 | ~133 [avatar contradiction clause 1445]
1447. $false [avatar sat refutation 339,344,349,354,359,369,379,551,566,582,1191,1195,1199,1205,1212,1341,1435,1446] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1097,thDist:-1097}
% SZS output end Proof for xreal_1__t135_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 87
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 112
% Generated clauses: 598
% Active clauses: 161
% Passive clauses: 249
% Final active clauses: 144
% Final passive clauses: 70
% Main loop iterations started: 173
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 74
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Forward subsumptions: 63
% 
% >>> Generating Inferences
% Binary resolution: 192
% Unit resulting resolution: 8
% Factoring: 1
% Forward superposition: 22
% Backward superposition: 6
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 36
% Split components: 89
% Sat splitting refutations: 16
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 184
% SAT solver unit clauses: 50
% SAT solver binary clauses: 68
% 
% Memory used [KB]: 10234
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% congruence closure: 0.001 s
% non unit clause index maintenance: 0.001 s
% forward subsumption: 0.002 s
% other: 0.006 s ( own 0.001 s ) 

