i: [1,179,0,2,1,1,0,0] t16_ordinal1
a: [3,281,0,2,1,179]
s: 281
i: [1,178,0,2,1,1,0,0] t16_ordinal1
a: [3,286,0,2,1,178]
s: 286
i: [1,181,0,3,1,1,0,0] t16_ordinal1
a: [3,271,0,3,1,181]
s: 271
i: [1,180,0,3,1,1,0,0] t16_ordinal1
a: [3,276,0,3,1,180]
s: 276
i: [1,182,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,291,0,2,1,182]
s: 291
i: [1,183,0,2,1,0,0,255] fc1_xboole_0
a: [3,296,0,2,1,183]
s: 296
i: [1,248,0,2,1,0,0,2] rc2_xboole_0
a: [3,313,0,2,1,248]
s: 313
i: [1,249,0,2,1,0,0,1] rc1_ordinal1
a: [3,318,0,2,1,249]
s: 318
i: [1,250,0,2,1,0,0,2] rc1_xboole_0
a: [3,323,0,2,1,250]
s: 323
i: [1,252,0,2,1,0,0,255] rc7_funct_1
a: [3,328,0,2,1,252]
s: 328
i: [1,251,0,2,1,0,0,255] rc7_funct_1
a: [3,333,0,2,1,251]
s: 333
i: [1,254,0,2,1,0,0,255] rc1_relat_1
a: [3,338,0,2,1,254]
s: 338
i: [1,253,0,2,1,0,0,255] rc1_relat_1
a: [3,343,0,2,1,253]
s: 343
i: [1,257,0,2,1,0,0,255] rc4_funct_1
a: [3,348,0,2,1,257]
s: 348
i: [1,256,0,2,1,0,0,255] rc4_funct_1
a: [3,353,0,2,1,256]
s: 353
i: [1,255,0,2,1,0,0,255] rc4_funct_1
a: [3,358,0,2,1,255]
s: 358
i: [1,258,0,2,1,0,0,255] rc2_relat_1
a: [3,363,0,2,1,258]
s: 363
i: [1,260,0,2,1,0,0,255] rc3_funct_1
a: [3,368,0,2,1,260]
s: 368
i: [1,259,0,2,1,0,0,255] rc3_funct_1
a: [3,373,0,2,1,259]
s: 373
i: [1,262,0,2,1,0,0,255] rc1_funct_1
a: [3,378,0,2,1,262]
s: 378
i: [1,261,0,2,1,0,0,255] rc1_funct_1
a: [3,383,0,2,1,261]
s: 383
i: [1,265,0,2,1,0,0,255] rc5_funct_1
a: [3,388,0,2,1,265]
s: 388
i: [1,264,0,2,1,0,0,255] rc5_funct_1
a: [3,393,0,2,1,264]
s: 393
i: [1,263,0,2,1,0,0,255] rc5_funct_1
a: [3,398,0,2,1,263]
s: 398
i: [1,267,0,2,1,0,0,255] rc2_funct_1
a: [3,403,0,2,1,267]
s: 403
i: [1,266,0,2,1,0,0,255] rc2_funct_1
a: [3,408,0,2,1,266]
s: 408
i: [1,185,0,3,1,0,0,2] fc1_subset_1
s: 185
i: [1,239,0,3,1,0,0,2] rc2_subset_1
s: 239
i: [1,242,0,3,1,0,0,2] reflexivity_r1_tarski
s: 242
i: [1,184,0,3,1,0,0,255] d2_xboole_0
a: [3,301,0,3,1,184]
s: 301
r: 291
i: [1,201,0,4,2,0,0,255] cc1_ordinal1
s: 201
d: [2,411,1,2,1,1,40,201,281]
a: [3,417,1,2,1,411]
s: 417
i: [1,202,0,4,2,0,0,255] cc1_ordinal1
s: 202
d: [2,412,1,2,1,1,40,201,286]
a: [3,422,1,2,1,412]
s: 422
i: [1,204,0,4,2,0,0,255] cc7_funct_1
s: 204
i: [1,205,0,4,2,0,0,255] cc1_funct_1
s: 205
d: [2,429,1,2,1,1,40,202,281]
a: [3,435,1,2,1,429]
s: 435
i: [1,206,0,4,2,0,0,255] cc1_relat_1
s: 206
d: [2,430,1,2,1,1,40,202,286]
a: [3,440,1,2,1,430]
s: 440
i: [1,237,0,4,1,0,0,2] existence_m1_subset_1
s: 237
i: [1,241,0,4,1,0,0,255] rc3_subset_1
s: 241
d: [2,413,1,2,1,1,40,201,318]
a: [3,427,1,2,1,413]
s: 427
i: [1,187,0,5,2,0,0,255] rc6_subset_1
s: 187
d: [2,431,1,2,1,1,40,202,318]
a: [3,445,1,2,1,431]
s: 445
i: [1,191,0,5,2,0,0,2] rc1_subset_1
s: 191
i: [1,193,0,5,2,0,0,255] rc5_subset_1
s: 193
d: [2,447,1,2,1,1,40,204,296]
a: [3,453,1,2,1,447]
s: 453
i: [1,194,0,5,2,0,0,255] rc5_subset_1
s: 194
d: [2,449,1,2,1,1,40,204,323]
a: [3,458,1,2,1,449]
s: 458
i: [1,207,0,5,2,0,0,255] t6_boole
s: 207
i: [1,234,0,5,2,0,0,1] t7_boole
s: 234
d: [2,460,1,2,1,1,40,205,296]
a: [3,466,1,2,1,460]
s: 466
i: [1,238,0,5,1,0,0,2] rc2_subset_1
s: 238
d: [2,462,1,2,1,1,40,205,323]
a: [3,471,1,2,1,462]
s: 471
i: [1,240,0,5,1,0,0,255] rc3_subset_1
s: 240
i: [1,244,0,7,3,0,0,1] reflexivity_r1_ordinal1
a: [3,308,0,5,2,244]
s: 308
d: [2,473,1,2,1,1,40,206,296]
a: [3,479,1,2,1,473]
s: 479
i: [1,196,0,6,2,0,0,255] rc4_subset_1
s: 196
d: [2,475,1,2,1,1,40,206,323]
a: [3,484,1,2,1,475]
s: 484
i: [1,211,0,6,3,0,0,255] cc2_ordinal1
s: 211
i: [1,227,0,6,2,0,0,255] d2_ordinal1
s: 227
d: [2,448,1,3,1,0,40,204,239]
s: 448
i: [1,228,0,6,2,0,0,255] d2_ordinal1
s: 228
d: [2,461,1,3,1,0,40,205,239]
s: 461
i: [1,229,0,6,2,0,0,1] t1_subset
s: 229
i: [1,243,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 243
d: [2,474,1,3,1,0,40,206,239]
s: 474
i: [1,186,0,7,2,0,0,255] rc6_subset_1
s: 186
d: [2,494,1,4,2,0,40,227,234]
s: 494
i: [1,190,0,7,2,0,0,2] rc1_subset_1
s: 190
i: [1,192,0,7,2,0,0,255] rc5_subset_1
s: 192
d: [2,497,2,2,1,1,40,494,296]
a: [3,503,2,2,1,497]
s: 503
i: [1,195,0,7,2,0,0,255] rc4_subset_1
s: 195
d: [2,499,2,2,1,1,40,494,323]
a: [3,508,2,2,1,499]
s: 508
i: [1,197,0,7,3,0,0,255] cc8_funct_1
s: 197
i: [1,198,0,7,3,0,0,255] cc8_funct_1
s: 198
d: [2,498,2,3,1,0,40,494,239]
s: 498
i: [1,231,0,7,2,0,0,2] t3_subset
s: 231
d: [2,530,1,4,1,0,40,231,238]
s: 530
i: [1,232,0,7,2,0,0,2] t3_subset
s: 232
i: [1,233,0,7,3,0,0,255] t8_boole
s: 233
d: [2,531,1,4,1,0,40,231,240]
s: 531
i: [1,199,0,8,3,0,0,255] cc9_funct_1
s: 199
d: [2,515,1,5,2,0,40,197,237]
s: 515
i: [1,200,0,8,3,0,0,255] cc5_subset_1
s: 200
i: [1,208,0,8,3,0,0,2] cc1_subset_1
s: 208
d: [2,522,1,5,2,0,40,198,237]
s: 522
i: [1,210,0,8,3,0,0,255] cc2_relat_1
s: 210
d: [2,529,1,5,1,0,40,231,237]
s: 529
i: [1,214,0,8,4,0,0,255] cc6_funct_1
s: 214
i: [1,219,0,8,4,0,0,255] cc4_funct_1
s: 219
d: [2,568,0,6,3,0,21,219,205]
d: [2,569,0,4,2,0,21,568,206]
s: 569
i: [1,222,0,8,4,0,0,255] cc5_funct_1
s: 222
d: [2,570,1,2,1,1,40,569,296]
a: [3,576,1,2,1,570]
s: 576
i: [1,226,0,8,3,0,0,255] d2_ordinal1
s: 226
i: [1,230,0,8,3,0,0,1] t2_subset
s: 230
d: [2,572,1,2,1,1,40,569,323]
a: [3,581,1,2,1,572]
s: 581
i: [1,209,0,9,3,0,0,255] cc4_subset_1
s: 209
d: [2,571,1,3,1,0,40,569,239]
s: 571
i: [1,236,0,9,3,0,0,1] t5_subset
s: 236
i: [1,225,0,10,4,0,0,255] cc3_funct_1
s: 225
d: [2,542,1,5,2,0,40,199,240]
s: 542
i: [1,235,0,10,3,0,0,1] t4_subset
s: 235
d: [2,549,1,5,2,0,40,200,238]
s: 549
i: [1,245,0,10,4,0,0,1] connectedness_r1_ordinal1
s: 245
i: [1,246,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 246
d: [2,633,1,7,3,1,40,245,276]
d: [2,636,1,5,2,2,21,633,286]
d: [2,637,1,3,1,3,21,636,281]
a: [3,641,1,3,1,637]
s: 641
i: [1,247,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 247
d: [2,647,2,7,3,1,40,641,246]
d: [2,648,2,5,2,2,21,647,286]
d: [2,649,2,3,1,3,21,648,281]
a: [3,653,2,3,1,649]
s: 653
i: [1,188,0,11,4,0,0,255] cc3_subset_1
s: 188
i: [1,189,0,11,4,0,0,255] cc2_subset_1
s: 189
d: [2,550,1,5,2,0,40,200,240]
s: 550
i: [1,203,0,13,5,0,0,255] t14_ordinal1
s: 203
d: [2,675,1,10,4,1,40,203,271]
d: [2,682,1,8,3,2,21,675,281]
d: [2,683,1,6,2,3,21,682,286]
a: [3,687,1,3,1,683]
s: 687
a: [3,690,1,3,1,683]
s: 690
d: [2,558,1,5,2,0,40,208,240]
s: 558
d: [2,697,2,2,1,1,40,687,234]
a: [3,712,2,2,1,697]
s: 712
d: [2,566,1,5,2,0,40,210,240]
s: 566
d: [2,694,2,5,2,1,40,687,226]
d: [2,698,2,3,1,2,21,694,417]
a: [3,702,2,3,1,698]
s: 702
d: [2,720,3,0,0,4,52,286,281,276,702,247]
e: 720
d: [2,721,3,7,3,1,40,702,247]
d: [2,723,3,4,2,2,21,721,276]
d: [2,724,3,2,1,3,21,723,281]
d: [2,725,3,0,0,4,21,724,286]
e: 725
f: 725,698,683,637,411,178,179,180,181
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal1__t16_ordinal1
% SZS output start Proof for ordinal1__t16_ordinal1
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_hidden(X1,X0) | r1_ordinal1(X0,X1)))) [input t16_ordinal1]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_hidden(X1,X0) | r1_ordinal1(X0,X1)))) [negated conjecture 1]
5. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
24. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1))) [input connectedness_r1_ordinal1]
25. ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r2_hidden(X1,X0) => r1_tarski(X1,X0))) [input d2_ordinal1]
51. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
54. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ~(~r2_hidden(X1,X0) & X0 != X1 & ~r2_hidden(X0,X1)))) [input t14_ordinal1]
78. ? [X0] : (? [X1] : ((~r2_hidden(X1,X0) & ~r1_ordinal1(X0,X1)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
79. ? [X0] : (? [X1] : (~r2_hidden(X1,X0) & ~r1_ordinal1(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 78]
91. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 5]
92. ! [X0] : (! [X1] : ((r2_hidden(X1,X0) | X0 = X1 | r2_hidden(X0,X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 54]
93. ! [X0] : (! [X1] : (r2_hidden(X1,X0) | X0 = X1 | r2_hidden(X0,X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 92]
117. ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0))) [ennf transformation 25]
129. ! [X0,X1] : ((r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 24]
130. ! [X0,X1] : (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 129]
131. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 51]
132. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 131]
133. ? [X0] : (? [X1] : (~r2_hidden(X1,X0) & ~r1_ordinal1(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (~r2_hidden(X1,sK0) & ~r1_ordinal1(sK0,X1) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
134. ? [X1] : (~r2_hidden(X1,sK0) & ~r1_ordinal1(sK0,X1) & v3_ordinal1(X1)) => (~r2_hidden(sK1,sK0) & ~r1_ordinal1(sK0,sK1) & v3_ordinal1(sK1)) [choice axiom]
135. (~r2_hidden(sK1,sK0) & ~r1_ordinal1(sK0,sK1) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 79,134,133]
144. ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0)) | ~v1_ordinal1(X0))) [nnf transformation 117]
145. ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0))) [rectify 144]
146. ! [X0] : (? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0)) => (~r1_tarski(sK6(X0),X0) & r2_hidden(sK6(X0),X0))) [choice axiom]
147. ! [X0] : ((v1_ordinal1(X0) | (~r1_tarski(sK6(X0),X0) & r2_hidden(sK6(X0),X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0))) [skolemisation 145,146]
155. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 132]
178. v3_ordinal1(sK0) [cnf transformation 135] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
179. v3_ordinal1(sK1) [cnf transformation 135] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
180. ~r1_ordinal1(sK0,sK1) [cnf transformation 135] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
181. ~r2_hidden(sK1,sK0) [cnf transformation 135] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
201. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 91] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
203. r2_hidden(X1,X0) | X0 = X1 | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 93] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
226. ~r2_hidden(X2,X0) | r1_tarski(X2,X0) | ~v1_ordinal1(X0) [cnf transformation 147] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
245. r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 130] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
247. ~r1_tarski(X0,X1) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 155] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
269. 1 <=> r2_hidden(sK1,sK0) [avatar definition]
271. ~r2_hidden(sK1,sK0) <- (~1) [avatar component clause 269] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
272. ~1 [avatar split clause 181,269]
274. 2 <=> r1_ordinal1(sK0,sK1) [avatar definition]
276. ~r1_ordinal1(sK0,sK1) <- (~2) [avatar component clause 274] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
277. ~2 [avatar split clause 180,274]
279. 3 <=> v3_ordinal1(sK1) [avatar definition]
281. v3_ordinal1(sK1) <- (3) [avatar component clause 279] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
282. 3 [avatar split clause 179,279]
284. 4 <=> v3_ordinal1(sK0) [avatar definition]
286. v3_ordinal1(sK0) <- (4) [avatar component clause 284] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
287. 4 [avatar split clause 178,284]
411. v1_ordinal1(sK1) <- (3) [resolution 201,281] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
415. 30 <=> v1_ordinal1(sK1) [avatar definition]
417. v1_ordinal1(sK1) <- (30) [avatar component clause 415] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
418. 30 | ~3 [avatar split clause 411,279,415]
633. r1_ordinal1(sK1,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK1) <- (~2) [resolution 245,276] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
636. r1_ordinal1(sK1,sK0) | ~v3_ordinal1(sK1) <- (~2, 4) [subsumption resolution 633,286] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
637. r1_ordinal1(sK1,sK0) <- (~2, 3, 4) [subsumption resolution 636,281] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
639. 46 <=> r1_ordinal1(sK1,sK0) [avatar definition]
642. 46 | 2 | ~3 | ~4 [avatar split clause 637,284,279,274,639]
675. sK0 = sK1 | r2_hidden(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~1) [resolution 203,271] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
682. sK0 = sK1 | r2_hidden(sK0,sK1) | ~v3_ordinal1(sK0) <- (~1, 3) [subsumption resolution 675,281] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
683. sK0 = sK1 | r2_hidden(sK0,sK1) <- (~1, 3, 4) [subsumption resolution 682,286] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
685. 48 <=> r2_hidden(sK0,sK1) [avatar definition]
687. r2_hidden(sK0,sK1) <- (48) [avatar component clause 685] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
689. 49 <=> sK0 = sK1 [avatar definition]
692. 48 | 49 | 1 | ~3 | ~4 [avatar split clause 683,284,279,269,689,685]
694. r1_tarski(sK0,sK1) | ~v1_ordinal1(sK1) <- (48) [resolution 687,226] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
698. r1_tarski(sK0,sK1) <- (30, 48) [subsumption resolution 694,417] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
700. 50 <=> r1_tarski(sK0,sK1) [avatar definition]
702. r1_tarski(sK0,sK1) <- (50) [avatar component clause 700] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
703. 50 | ~30 | ~48 [avatar split clause 698,685,415,700]
714. sK0 != sK1 | r1_ordinal1(sK0,sK1) | ~r1_ordinal1(sK1,sK0) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
721. r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (50) [resolution 702,247] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
723. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~2, 50) [subsumption resolution 721,276] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
724. ~v3_ordinal1(sK0) <- (~2, 3, 50) [subsumption resolution 723,281] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
725. $false <- (~2, 3, 4, 50) [subsumption resolution 724,286] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
726. 2 | ~3 | ~4 | ~50 [avatar contradiction clause 725]
727. $false [avatar sat refutation 272,277,282,287,418,642,692,703,714,726] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:241,thDist:-241}
% SZS output end Proof for ordinal1__t16_ordinal1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 61
% >>> Preprocessing
% Introduced skolems: 21
% Pure predicates: 3
% 
% >>> Saturation
% Initial clauses: 90
% Generated clauses: 341
% Active clauses: 116
% Passive clauses: 209
% Final active clauses: 115
% Final passive clauses: 90
% Main loop iterations started: 119
% 
% >>> Simplifying Inferences
% Duplicate literals: 10
% Fw subsumption resolutions: 22
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 24
% Equational tautologies: 1
% Forward subsumptions: 28
% 
% >>> Generating Inferences
% Binary resolution: 159
% Unit resulting resolution: 1
% Factoring: 2
% Forward superposition: 4
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 55
% SAT solver unit clauses: 27
% SAT solver binary clauses: 19
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% forward subsumption resolution: 0.001 s
% resolution: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

