i: [1,70,0,2,1,1,0,0] t5_procal_1
a: [3,114,0,2,1,70]
s: 114
i: [1,72,0,6,1,1,0,0] t5_procal_1
a: [3,104,0,6,1,72]
s: 104
i: [1,71,0,6,1,1,0,0] t5_procal_1
a: [3,109,0,6,1,71]
s: 109
i: [1,95,0,2,1,0,0,1] existence_m1_qc_lang1
a: [3,119,0,2,1,95]
s: 119
i: [1,73,0,15,1,1,0,0] t5_procal_1
a: [3,99,0,15,1,73]
s: 99
i: [1,80,0,4,1,0,0,2] existence_m1_subset_1
s: 80
i: [1,74,0,5,2,0,0,3] dt_k9_qc_lang1
s: 74
i: [1,75,0,5,2,0,0,3] fc2_cqc_lang
s: 75
i: [1,81,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 81
i: [1,76,0,8,2,0,0,1] dt_k4_cqc_the1
s: 76
i: [1,77,0,8,2,0,0,2] dt_k3_cqc_lang
s: 77
i: [1,86,0,12,3,0,0,255] dt_k13_qc_lang1
s: 86
i: [1,87,0,13,3,0,0,255] redefinition_k6_cqc_lang
s: 87
i: [1,85,0,14,4,0,0,2] existence_m2_subset_1
s: 85
i: [1,88,0,14,3,0,0,1] dt_k6_cqc_lang
s: 88
i: [1,78,0,15,3,0,0,1] t4_lukasi_1
s: 78
i: [1,82,0,15,5,0,0,2] dt_m2_subset_1
s: 82
i: [1,83,0,15,5,0,0,2] redefinition_m2_subset_1
s: 83
d: [2,124,1,18,5,0,40,82,77]
d: [2,127,1,15,4,0,21,124,75]
d: [2,128,1,12,3,0,21,127,74]
s: 128
i: [1,84,0,15,5,0,0,2] redefinition_m2_subset_1
s: 84
i: [1,89,0,17,4,0,0,255] dt_k2_qc_lang2
s: 89
d: [2,136,2,6,2,1,40,128,104]
d: [2,141,2,4,1,2,21,136,114]
a: [3,145,2,4,1,141]
s: 145
i: [1,90,0,17,4,0,0,255] dt_k3_qc_lang2
s: 90
d: [2,137,2,6,2,1,40,128,109]
d: [2,147,2,4,1,2,21,137,114]
a: [3,151,2,4,1,147]
s: 151
i: [1,91,0,19,4,0,0,255] redefinition_k8_cqc_lang
s: 91
i: [1,92,0,19,4,0,0,255] redefinition_k9_cqc_lang
s: 92
d: [2,139,2,22,5,0,40,128,85]
d: [2,153,2,16,4,0,21,139,77]
d: [2,154,2,13,3,0,21,153,75]
d: [2,155,2,10,2,0,21,154,74]
s: 155
i: [1,93,0,19,4,0,0,1] dt_k9_cqc_lang
s: 93
d: [2,130,1,18,5,0,40,83,77]
d: [2,133,1,15,4,0,21,130,75]
d: [2,134,1,12,3,0,21,133,74]
s: 134
i: [1,94,0,19,4,0,0,1] dt_k8_cqc_lang
s: 94
i: [1,79,0,25,4,0,0,1] l2_procal_1
s: 79
d: [2,168,2,6,2,1,40,134,104]
d: [2,175,2,4,1,2,21,168,114]
a: [3,179,2,4,1,175]
s: 179
d: [2,194,1,27,4,1,44,99,79]
d: [2,199,1,21,3,2,21,194,104]
d: [2,200,1,15,2,3,21,199,109]
d: [2,201,1,13,1,4,21,200,114]
a: [3,205,1,13,1,201]
s: 205
d: [2,169,2,6,2,1,40,134,109]
d: [2,181,2,4,1,2,21,169,114]
a: [3,185,2,4,1,181]
s: 185
d: [2,156,1,18,5,0,40,84,77]
d: [2,159,1,15,4,0,21,156,75]
d: [2,160,1,12,3,0,21,159,74]
s: 160
d: [2,123,1,15,3,0,40,78,81]
s: 123
d: [2,208,2,11,2,1,40,205,78]
d: [2,209,2,9,1,2,21,208,114]
a: [3,213,2,9,1,209]
s: 213
d: [2,221,3,0,0,4,52,114,185,179,213,93]
e: 221
d: [2,222,3,10,3,1,40,213,93]
d: [2,225,3,6,2,2,21,222,179]
d: [2,226,3,2,1,3,21,225,185]
d: [2,227,3,0,0,4,21,226,114]
e: 227
f: 227,209,201,181,175,70,71,72,73
% Refutation found. Thanks to Tanya!
% SZS status Theorem for procal_1__t5_procal_1
% SZS output start Proof for procal_1__t5_procal_1
1. ! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0)) => ! [X2] : (m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0)) => r2_hidden(k8_cqc_lang(X0,k9_cqc_lang(X0,X1,X2),k8_cqc_lang(X0,k6_cqc_lang(X0,X1),X2)),k4_cqc_the1(X0))))) [input t5_procal_1]
2. ~! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0)) => ! [X2] : (m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0)) => r2_hidden(k8_cqc_lang(X0,k9_cqc_lang(X0,X1,X2),k8_cqc_lang(X0,k6_cqc_lang(X0,X1),X2)),k4_cqc_the1(X0))))) [negated conjecture 1]
7. ! [X0] : (m1_qc_lang1(X0) => m1_subset_1(k3_cqc_lang(X0),k1_zfmisc_1(k9_qc_lang1(X0)))) [input dt_k3_cqc_lang]
12. ! [X0,X1,X2] : ((m1_subset_1(X2,k3_cqc_lang(X0)) & m1_subset_1(X1,k3_cqc_lang(X0)) & m1_qc_lang1(X0)) => m2_subset_1(k9_cqc_lang(X0,X1,X2),k9_qc_lang1(X0),k3_cqc_lang(X0))) [input dt_k9_cqc_lang]
13. ! [X0] : (m1_qc_lang1(X0) => ~v1_xboole_0(k9_qc_lang1(X0))) [input dt_k9_qc_lang1]
20. ! [X0] : (m1_qc_lang1(X0) => ~v1_xboole_0(k3_cqc_lang(X0))) [input fc2_cqc_lang]
21. ! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0)) => ! [X2] : (m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0)) => k9_cqc_lang(X0,X1,X2) = k8_cqc_lang(X0,k6_cqc_lang(X0,X1),X2)))) [input l2_procal_1]
25. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ! [X2] : (m2_subset_1(X2,X0,X1) <=> m1_subset_1(X2,X1))) [input redefinition_m2_subset_1]
26. ! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0)) => r2_hidden(k8_cqc_lang(X0,X1,X1),k4_cqc_the1(X0)))) [input t4_lukasi_1]
27. ? [X0] : (? [X1] : (? [X2] : (~r2_hidden(k8_cqc_lang(X0,k9_cqc_lang(X0,X1,X2),k8_cqc_lang(X0,k6_cqc_lang(X0,X1),X2)),k4_cqc_the1(X0)) & m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0))) & m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0))) & m1_qc_lang1(X0)) [ennf transformation 2]
28. ! [X0] : (~v1_xboole_0(k9_qc_lang1(X0)) | ~m1_qc_lang1(X0)) [ennf transformation 13]
29. ! [X0] : (~v1_xboole_0(k3_cqc_lang(X0)) | ~m1_qc_lang1(X0)) [ennf transformation 20]
31. ! [X0] : (m1_subset_1(k3_cqc_lang(X0),k1_zfmisc_1(k9_qc_lang1(X0))) | ~m1_qc_lang1(X0)) [ennf transformation 7]
32. ! [X0] : (! [X1] : (r2_hidden(k8_cqc_lang(X0,X1,X1),k4_cqc_the1(X0)) | ~m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0))) | ~m1_qc_lang1(X0)) [ennf transformation 26]
33. ! [X0] : (! [X1] : (! [X2] : (k9_cqc_lang(X0,X1,X2) = k8_cqc_lang(X0,k6_cqc_lang(X0,X1),X2) | ~m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0))) | ~m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0))) | ~m1_qc_lang1(X0)) [ennf transformation 21]
37. ! [X0,X1] : (! [X2] : (m2_subset_1(X2,X0,X1) <=> m1_subset_1(X2,X1)) | (~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 25]
38. ! [X0,X1] : (! [X2] : (m2_subset_1(X2,X0,X1) <=> m1_subset_1(X2,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 37]
55. ! [X0,X1,X2] : (m2_subset_1(k9_cqc_lang(X0,X1,X2),k9_qc_lang1(X0),k3_cqc_lang(X0)) | (~m1_subset_1(X2,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0))) [ennf transformation 12]
56. ! [X0,X1,X2] : (m2_subset_1(k9_cqc_lang(X0,X1,X2),k9_qc_lang1(X0),k3_cqc_lang(X0)) | ~m1_subset_1(X2,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0)) [flattening 55]
59. ? [X0] : (? [X1] : (? [X2] : (~r2_hidden(k8_cqc_lang(X0,k9_cqc_lang(X0,X1,X2),k8_cqc_lang(X0,k6_cqc_lang(X0,X1),X2)),k4_cqc_the1(X0)) & m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0))) & m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0))) & m1_qc_lang1(X0)) => (? [X1] : (? [X2] : (~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,X1,X2),k8_cqc_lang(sK0,k6_cqc_lang(sK0,X1),X2)),k4_cqc_the1(sK0)) & m2_subset_1(X2,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) & m2_subset_1(X1,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) & m1_qc_lang1(sK0)) [choice axiom]
60. ? [X1] : (? [X2] : (~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,X1,X2),k8_cqc_lang(sK0,k6_cqc_lang(sK0,X1),X2)),k4_cqc_the1(sK0)) & m2_subset_1(X2,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) & m2_subset_1(X1,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) => (? [X2] : (~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,X2),k8_cqc_lang(sK0,k6_cqc_lang(sK0,sK1),X2)),k4_cqc_the1(sK0)) & m2_subset_1(X2,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) & m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) [choice axiom]
61. ? [X2] : (~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,X2),k8_cqc_lang(sK0,k6_cqc_lang(sK0,sK1),X2)),k4_cqc_the1(sK0)) & m2_subset_1(X2,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) => (~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k8_cqc_lang(sK0,k6_cqc_lang(sK0,sK1),sK2)),k4_cqc_the1(sK0)) & m2_subset_1(sK2,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) [choice axiom]
62. ((~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k8_cqc_lang(sK0,k6_cqc_lang(sK0,sK1),sK2)),k4_cqc_the1(sK0)) & m2_subset_1(sK2,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) & m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) & m1_qc_lang1(sK0) [skolemisation 27,61,60,59]
65. ! [X0,X1] : (! [X2] : ((m2_subset_1(X2,X0,X1) | ~m1_subset_1(X2,X1)) & (m1_subset_1(X2,X1) | ~m2_subset_1(X2,X0,X1))) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [nnf transformation 38]
70. m1_qc_lang1(sK0) [cnf transformation 62] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
71. m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) [cnf transformation 62] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
72. m2_subset_1(sK2,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) [cnf transformation 62] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
73. ~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k8_cqc_lang(sK0,k6_cqc_lang(sK0,sK1),sK2)),k4_cqc_the1(sK0)) [cnf transformation 62] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
74. ~v1_xboole_0(k9_qc_lang1(X0)) | ~m1_qc_lang1(X0) [cnf transformation 28] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
75. ~v1_xboole_0(k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [cnf transformation 29] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
77. m1_subset_1(k3_cqc_lang(X0),k1_zfmisc_1(k9_qc_lang1(X0))) | ~m1_qc_lang1(X0) [cnf transformation 31] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
78. r2_hidden(k8_cqc_lang(X0,X1,X1),k4_cqc_the1(X0)) | ~m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [cnf transformation 32] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. k9_cqc_lang(X0,X1,X2) = k8_cqc_lang(X0,k6_cqc_lang(X0,X1),X2) | ~m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0)) | ~m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [cnf transformation 33] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
83. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | ~m2_subset_1(X2,X0,X1) | m1_subset_1(X2,X1) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 65] {a:0,w:15,wCS:6000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
93. m2_subset_1(k9_cqc_lang(X0,X1,X2),k9_qc_lang1(X0),k3_cqc_lang(X0)) | ~m1_subset_1(X2,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [cnf transformation 56] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
97. 1 <=> r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k8_cqc_lang(sK0,k6_cqc_lang(sK0,sK1),sK2)),k4_cqc_the1(sK0)) [avatar definition]
99. ~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k8_cqc_lang(sK0,k6_cqc_lang(sK0,sK1),sK2)),k4_cqc_the1(sK0)) <- (~1) [avatar component clause 97] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
100. ~1 [avatar split clause 73,97]
102. 2 <=> m2_subset_1(sK2,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) [avatar definition]
104. m2_subset_1(sK2,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) <- (2) [avatar component clause 102] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
105. 2 [avatar split clause 72,102]
107. 3 <=> m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) [avatar definition]
109. m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) <- (3) [avatar component clause 107] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
110. 3 [avatar split clause 71,107]
112. 4 <=> m1_qc_lang1(sK0) [avatar definition]
114. m1_qc_lang1(sK0) <- (4) [avatar component clause 112] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
115. 4 [avatar split clause 70,112]
130. ~m2_subset_1(X0,k9_qc_lang1(X1),k3_cqc_lang(X1)) | m1_subset_1(X0,k3_cqc_lang(X1)) | v1_xboole_0(k3_cqc_lang(X1)) | v1_xboole_0(k9_qc_lang1(X1)) | ~m1_qc_lang1(X1) [resolution 83,77] {a:1,w:18,wCS:7200,sine:2,thAx:0,allAx:2,thDist:-2}
133. ~m2_subset_1(X0,k9_qc_lang1(X1),k3_cqc_lang(X1)) | m1_subset_1(X0,k3_cqc_lang(X1)) | v1_xboole_0(k9_qc_lang1(X1)) | ~m1_qc_lang1(X1) [subsumption resolution 130,75] {a:1,w:15,wCS:6000,sine:2,thAx:0,allAx:2,thDist:-2}
134. ~m2_subset_1(X0,k9_qc_lang1(X1),k3_cqc_lang(X1)) | m1_subset_1(X0,k3_cqc_lang(X1)) | ~m1_qc_lang1(X1) [subsumption resolution 133,74] {a:1,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
168. m1_subset_1(sK2,k3_cqc_lang(sK0)) | ~m1_qc_lang1(sK0) <- (2) [resolution 134,104] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
169. m1_subset_1(sK1,k3_cqc_lang(sK0)) | ~m1_qc_lang1(sK0) <- (3) [resolution 134,109] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
175. m1_subset_1(sK2,k3_cqc_lang(sK0)) <- (2, 4) [subsumption resolution 168,114] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
177. 8 <=> m1_subset_1(sK2,k3_cqc_lang(sK0)) [avatar definition]
179. m1_subset_1(sK2,k3_cqc_lang(sK0)) <- (8) [avatar component clause 177] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
180. 8 | ~2 | ~4 [avatar split clause 175,112,102,177]
181. m1_subset_1(sK1,k3_cqc_lang(sK0)) <- (3, 4) [subsumption resolution 169,114] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
183. 9 <=> m1_subset_1(sK1,k3_cqc_lang(sK0)) [avatar definition]
185. m1_subset_1(sK1,k3_cqc_lang(sK0)) <- (9) [avatar component clause 183] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
186. 9 | ~3 | ~4 [avatar split clause 181,112,107,183]
194. ~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k9_cqc_lang(sK0,sK1,sK2)),k4_cqc_the1(sK0)) | ~m2_subset_1(sK2,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) | ~m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) | ~m1_qc_lang1(sK0) <- (~1) [superposition 99,79] {a:1,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
199. ~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k9_cqc_lang(sK0,sK1,sK2)),k4_cqc_the1(sK0)) | ~m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) | ~m1_qc_lang1(sK0) <- (~1, 2) [subsumption resolution 194,104] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
200. ~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k9_cqc_lang(sK0,sK1,sK2)),k4_cqc_the1(sK0)) | ~m1_qc_lang1(sK0) <- (~1, 2, 3) [subsumption resolution 199,109] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
201. ~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k9_cqc_lang(sK0,sK1,sK2)),k4_cqc_the1(sK0)) <- (~1, 2, 3, 4) [subsumption resolution 200,114] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
203. 10 <=> r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k9_cqc_lang(sK0,sK1,sK2)),k4_cqc_the1(sK0)) [avatar definition]
205. ~r2_hidden(k8_cqc_lang(sK0,k9_cqc_lang(sK0,sK1,sK2),k9_cqc_lang(sK0,sK1,sK2)),k4_cqc_the1(sK0)) <- (~10) [avatar component clause 203] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
206. ~10 | 1 | ~2 | ~3 | ~4 [avatar split clause 201,112,107,102,97,203]
208. ~m2_subset_1(k9_cqc_lang(sK0,sK1,sK2),k9_qc_lang1(sK0),k3_cqc_lang(sK0)) | ~m1_qc_lang1(sK0) <- (~10) [resolution 205,78] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
209. ~m2_subset_1(k9_cqc_lang(sK0,sK1,sK2),k9_qc_lang1(sK0),k3_cqc_lang(sK0)) <- (4, ~10) [subsumption resolution 208,114] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
211. 11 <=> m2_subset_1(k9_cqc_lang(sK0,sK1,sK2),k9_qc_lang1(sK0),k3_cqc_lang(sK0)) [avatar definition]
213. ~m2_subset_1(k9_cqc_lang(sK0,sK1,sK2),k9_qc_lang1(sK0),k3_cqc_lang(sK0)) <- (~11) [avatar component clause 211] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
214. ~11 | ~4 | 10 [avatar split clause 209,203,112,211]
222. ~m1_subset_1(sK2,k3_cqc_lang(sK0)) | ~m1_subset_1(sK1,k3_cqc_lang(sK0)) | ~m1_qc_lang1(sK0) <- (~11) [resolution 213,93] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
225. ~m1_subset_1(sK1,k3_cqc_lang(sK0)) | ~m1_qc_lang1(sK0) <- (8, ~11) [subsumption resolution 222,179] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
226. ~m1_qc_lang1(sK0) <- (8, 9, ~11) [subsumption resolution 225,185] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
227. $false <- (4, 8, 9, ~11) [subsumption resolution 226,114] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
228. ~4 | ~8 | ~9 | 11 [avatar contradiction clause 227]
230. $false [avatar sat refutation 100,105,110,115,180,186,206,214,228] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:108,thDist:-108}
% SZS output end Proof for procal_1__t5_procal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 114
% Active clauses: 37
% Passive clauses: 57
% Final active clauses: 37
% Final passive clauses: 20
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Duplicate literals: 16
% Fw subsumption resolutions: 28
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% 
% >>> Generating Inferences
% Binary resolution: 27
% Unit resulting resolution: 1
% Forward superposition: 1
% Backward superposition: 6
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

