i: [1,65609,0,51,16,1,0,0] t4_gate_3
a: [3,65629,0,2,1,65609]
s: 65629
a: [3,65637,0,2,1,65609]
s: 65637
a: [3,65644,0,2,1,65609]
s: 65644
a: [3,65652,0,2,1,65609]
s: 65652
a: [3,65661,0,2,1,65609]
s: 65661
a: [3,65669,0,2,1,65609]
s: 65669
a: [3,65676,0,2,1,65609]
s: 65676
a: [3,65684,0,2,1,65609]
s: 65684
i: [1,73,0,6,2,1,0,0] t4_gate_3
a: [3,65949,0,2,1,73]
s: 65949
i: [1,71,0,6,2,1,0,0] t4_gate_3
a: [3,65958,0,2,1,71]
s: 65958
i: [1,69,0,7,2,1,0,0] t4_gate_3
a: [3,65968,0,2,1,69]
s: 65968
i: [1,51,0,7,2,1,0,0] t4_gate_3
a: [3,66026,0,2,1,51]
s: 66026
i: [1,49,0,8,2,1,0,0] t4_gate_3
a: [3,66036,0,2,1,49]
s: 66036
i: [1,47,0,8,2,1,0,0] t4_gate_3
a: [3,66046,0,2,1,47]
s: 66046
i: [1,45,0,9,2,1,0,0] t4_gate_3
a: [3,66057,0,2,1,45]
s: 66057
i: [1,43,0,8,2,1,0,0] t4_gate_3
a: [3,66066,0,2,1,43]
s: 66066
i: [1,41,0,9,2,1,0,0] t4_gate_3
a: [3,66077,0,2,1,41]
s: 66077
i: [1,39,0,9,2,1,0,0] t4_gate_3
a: [3,66086,0,2,1,39]
s: 66086
i: [1,37,0,10,2,1,0,0] t4_gate_3
a: [3,66097,0,2,1,37]
s: 66097
a: [3,65641,0,4,1,65609]
s: 65641
a: [3,65944,0,4,1,73]
s: 65944
a: [3,65955,0,4,1,71]
s: 65955
a: [3,65965,0,5,1,69]
s: 65965
i: [1,67,0,7,2,1,0,0] t4_gate_3
a: [3,65974,0,5,1,67]
s: 65974
a: [3,66023,0,5,1,51]
s: 66023
i: [1,65,0,8,2,1,0,0] t4_gate_3
a: [3,65980,0,6,1,65]
s: 65980
i: [1,63,0,8,2,1,0,0] t4_gate_3
a: [3,65986,0,6,1,63]
s: 65986
i: [1,59,0,8,2,1,0,0] t4_gate_3
a: [3,65999,0,6,1,59]
s: 65999
a: [3,66033,0,6,1,49]
s: 66033
a: [3,66043,0,6,1,47]
s: 66043
a: [3,66063,0,6,1,43]
s: 66063
i: [1,61,0,9,2,1,0,0] t4_gate_3
a: [3,65993,0,7,1,61]
s: 65993
i: [1,57,0,9,2,1,0,0] t4_gate_3
a: [3,66004,0,7,1,57]
s: 66004
i: [1,55,0,9,2,1,0,0] t4_gate_3
a: [3,66011,0,7,1,55]
s: 66011
a: [3,66052,0,7,1,45]
s: 66052
a: [3,66072,0,7,1,41]
s: 66072
a: [3,66083,0,7,1,39]
s: 66083
i: [1,53,0,10,2,1,0,0] t4_gate_3
a: [3,66017,0,8,1,53]
s: 66017
a: [3,66092,0,8,1,37]
s: 66092
i: [1,65610,0,5,2,0,0,1] fc2_gate_1
s: 65610
i: [1,65611,0,5,2,0,0,1] fc1_gate_1
s: 65611
i: [1,65614,0,6,2,0,0,1] fc5_gate_1
s: 65614
i: [1,65615,0,6,2,0,0,1] fc4_gate_1
s: 65615
d: [2,66101,1,2,1,1,40,65615,65944]
a: [3,66105,1,2,1,66101]
s: 66105
i: [1,65612,0,7,1,0,0,255] commutativity_k3_gate_1
s: 65612
i: [1,65613,0,7,1,0,0,255] commutativity_k2_gate_1
s: 65613
r: 65965
r: 65968
a: [3,65969,0,2,1,69]
s: 65969
d: [2,66109,0,5,1,1,23,65965,65613]
a: [3,66114,0,5,1,66109]
s: 66114
d: [2,66119,1,2,1,1,40,66114,65615]
a: [3,66123,1,2,1,66119]
s: 66123
i: [1,65618,0,7,2,0,0,1] fc18_gate_1
s: 65618
i: [1,65619,0,7,2,0,0,1] fc19_gate_1
s: 65619
d: [2,66148,1,0,0,2,52,65958,65974,65619]
e: 66148
d: [2,66149,1,0,0,2,52,65958,65980,65619]
e: 66149
d: [2,66150,1,0,0,2,52,65958,66004,65619]
e: 66150
d: [2,66151,1,2,1,1,40,65619,65974]
d: [2,66161,1,0,0,2,21,66151,65958]
e: 66161
d: [2,66152,1,2,1,1,40,65619,65980]
d: [2,66163,1,0,0,2,21,66152,65958]
e: 66163
d: [2,66154,1,2,1,1,40,65619,66004]
d: [2,66165,1,0,0,2,21,66154,65958]
e: 66165
r: 65955
r: 65958
a: [3,65959,0,2,1,71]
s: 65959
a: [3,65954,0,4,1,71]
s: 65954
d: [2,66131,1,2,1,1,40,65618,66052]
a: [3,66146,1,2,1,66131]
s: 66146
i: [1,65620,0,7,2,0,0,1] fc20_gate_1
s: 65620
i: [1,65616,0,8,3,0,0,1] fc3_gate_1
s: 65616
d: [2,66181,1,2,1,1,40,65954,65615]
a: [3,66185,1,2,1,66181]
s: 66185
i: [1,65617,0,8,3,0,0,1] fc6_gate_1
s: 65617
d: [2,66128,1,3,1,1,40,65618,66004]
a: [3,66135,1,3,1,66128]
s: 66135
d: [2,66214,2,0,0,2,52,65949,66135,65610]
e: 66214
d: [2,66215,2,2,1,1,40,66135,65610]
d: [2,66217,2,0,0,2,21,66215,65949]
e: 66217
r: 65676
r: 66004
r: 66135
a: [3,65677,0,2,1,65609]
s: 65677
a: [3,66005,0,7,1,57]
s: 66005
d: [2,66129,1,3,1,1,40,65618,66072]
a: [3,66140,1,3,1,66129]
s: 66140
d: [2,66219,2,0,0,2,52,66146,66140,65610]
e: 66219
d: [2,66220,2,2,1,1,40,66140,65610]
d: [2,66222,2,0,0,2,21,66220,66146]
e: 66222
r: 66072
r: 66077
r: 66092
r: 66097
r: 66140
a: [3,66076,0,2,1,41]
s: 66076
a: [3,66096,0,2,1,37]
s: 66096
a: [3,66073,0,7,1,41]
s: 66073
a: [3,66093,0,8,1,37]
s: 66093
d: [2,66157,1,3,1,1,40,65619,66052]
a: [3,66173,1,3,1,66157]
s: 66173
d: [2,66224,2,0,0,2,52,66105,66173,65610]
e: 66224
d: [2,66225,2,2,1,1,40,66173,65610]
d: [2,66227,2,0,0,2,21,66225,66105]
e: 66227
r: 66052
r: 66057
r: 66173
a: [3,66056,0,2,1,45]
s: 66056
d: [2,66153,1,3,1,1,40,65619,65986]
a: [3,66179,1,3,1,66153]
s: 66179
d: [2,66229,2,0,0,2,52,65959,66179,65610]
e: 66229
d: [2,66230,2,2,1,1,40,66179,65610]
d: [2,66232,2,0,0,2,21,66230,65959]
e: 66232
r: 65644
r: 65652
r: 65954
r: 65959
r: 65974
r: 65980
a: [3,65645,0,2,1,65609]
s: 65645
a: [3,65653,0,2,1,65609]
s: 65653
s: 65958
a: [3,65632,0,4,1,65609]
s: 65632
d: [2,66235,1,0,0,2,52,66086,65632,65615]
e: 66235
d: [2,66236,1,2,1,1,40,65632,65615]
d: [2,66238,1,0,0,2,21,66236,66086]
e: 66238
r: 66083
r: 66086
a: [3,66087,0,2,1,39]
s: 66087
s: 65955
d: [2,66241,2,0,0,3,52,66185,66123,65955,65616]
e: 66241
d: [2,66242,1,4,2,1,40,65955,65616]
d: [2,66244,1,2,1,2,21,66242,66123]
d: [2,66245,1,0,0,3,21,66244,66185]
e: 66245
r: 65969
r: 66114
r: 66123
r: 66185
s: 65969
s: 66123
a: [3,65975,0,5,1,67]
s: 65975
a: [3,66184,1,2,1,66181]
s: 66184
s: 66114
a: [3,65981,0,6,1,65]
s: 65981
d: [2,66188,1,3,1,1,40,65620,65980]
a: [3,66197,1,3,1,66188]
s: 66197
d: [2,66256,2,0,0,2,52,65969,66197,65610]
e: 66256
d: [2,66257,2,2,1,1,40,66197,65610]
d: [2,66259,2,0,0,2,21,66257,65969]
e: 66259
r: 66197
a: [3,66053,0,7,1,45]
s: 66053
d: [2,66191,1,3,1,1,40,65620,66072]
a: [3,66203,1,3,1,66191]
s: 66203
a: [3,66082,0,7,1,39]
s: 66082
d: [2,66262,2,0,0,2,52,66184,66082,65620]
e: 66262
d: [2,66263,1,2,1,1,40,66082,65620]
d: [2,66267,1,0,0,2,21,66263,66184]
e: 66267
r: 65684
r: 65969
r: 65986
r: 66114
r: 66123
r: 66184
r: 65632
r: 66082
r: 66087
r: 66203
a: [3,65685,0,2,1,65609]
s: 65685
s: 65968
s: 66086
a: [3,66122,1,2,1,66119]
s: 66122
a: [3,65633,0,4,1,65609]
s: 65633
a: [3,65648,0,4,1,65609]
s: 65648
d: [2,66276,1,0,0,2,52,66026,65648,65615]
e: 66276
d: [2,66277,1,2,1,1,40,65648,65615]
d: [2,66279,1,0,0,2,21,66277,66026]
e: 66279
r: 66023
r: 66026
a: [3,66027,0,2,1,51]
s: 66027
a: [3,66022,0,5,1,51]
s: 66022
s: 66185
a: [3,66115,0,5,1,66109]
s: 66115
a: [3,65987,0,6,1,63]
s: 65987
a: [3,66202,1,3,1,66191]
s: 66202
s: 66083
d: [2,66107,1,6,2,0,44,65614,65612]
s: 66107
i: [1,65621,0,11,4,0,0,1] fc17_gate_1
s: 65621
a: [3,66139,1,3,1,66129]
s: 66139
a: [3,66172,1,3,1,66157]
s: 66172
a: [3,66196,1,3,1,66188]
s: 66196
d: [2,66324,2,0,0,2,52,65968,66196,65611]
e: 66324
d: [2,66325,2,2,1,1,40,66196,65611]
d: [2,66327,2,0,0,2,21,66325,65968]
e: 66327
r: 65661
r: 65993
r: 66196
a: [3,65660,0,2,1,65609]
s: 65660
a: [3,65992,0,7,1,61]
s: 65992
a: [3,66134,1,3,1,66128]
s: 66134
s: 66197
d: [2,66108,1,6,2,0,44,65614,65612]
s: 66108
d: [2,66110,1,6,2,0,44,65615,65613]
s: 66110
d: [2,66335,2,0,0,2,52,66122,65944,66110]
e: 66335
d: [2,66336,2,0,0,2,52,66122,65648,66110]
e: 66336
d: [2,66338,2,2,1,1,40,66110,65648]
d: [2,66344,2,0,0,2,21,66338,66122]
e: 66344
d: [2,66339,2,2,1,1,40,66110,65944]
d: [2,66346,2,0,0,2,21,66339,66122]
e: 66346
r: 65648
r: 65944
r: 65949
r: 65660
r: 65669
r: 65992
r: 66017
r: 66134
s: 65661
a: [3,65668,0,2,1,65609]
s: 65668
a: [3,65948,0,2,1,73]
s: 65948
a: [3,65657,0,4,1,65609]
s: 65657
s: 66135
a: [3,65672,0,4,1,65609]
s: 65672
d: [2,66352,2,0,0,2,52,66122,65672,66110]
e: 66352
d: [2,66353,1,0,0,2,52,66046,65672,65615]
e: 66353
d: [2,66354,2,2,1,1,40,65672,66110]
d: [2,66358,2,0,0,2,21,66354,66122]
e: 66358
d: [2,66355,1,2,1,1,40,65672,65615]
d: [2,66360,1,0,0,2,21,66355,66046]
e: 66360
r: 65672
a: [3,65945,0,4,1,73]
s: 65945
s: 65993
a: [3,65665,0,7,1,65609]
d: [2,66362,0,7,1,1,22,65665,65612]
a: [3,66366,0,7,1,66362]
s: 66366
d: [2,66370,2,3,1,1,40,66366,66107]
a: [3,66375,2,3,1,66370]
s: 66375
d: [2,66382,3,0,0,2,52,66122,66375,65611]
e: 66382
d: [2,66383,3,2,1,1,40,66375,65611]
d: [2,66385,3,0,0,2,21,66383,66122]
e: 66385
r: 66366
r: 66375
a: [3,65680,0,4,1,65609]
s: 65680
d: [2,66389,2,0,0,2,52,66122,65680,66110]
e: 66389
d: [2,66390,1,0,0,2,52,66076,65680,65615]
e: 66390
d: [2,66391,2,2,1,1,40,65680,66110]
d: [2,66395,2,0,0,2,21,66391,66122]
e: 66395
d: [2,66392,1,2,1,1,40,65680,65615]
d: [2,66397,1,0,0,2,21,66392,66076]
e: 66397
r: 65641
r: 65680
a: [3,65640,0,4,1,65609]
s: 65640
d: [2,66399,2,0,0,2,52,66122,65640,66110]
e: 66399
d: [2,66400,1,0,0,2,52,66066,65640,65615]
e: 66400
d: [2,66401,2,2,1,1,40,65640,66110]
d: [2,66405,2,0,0,2,21,66401,66122]
e: 66405
d: [2,66402,1,2,1,1,40,65640,65615]
d: [2,66407,1,0,0,2,21,66402,66066]
e: 66407
r: 65640
r: 65657
a: [3,65656,0,4,1,65609]
s: 65656
d: [2,66409,2,0,0,2,52,66122,65656,66110]
e: 66409
d: [2,66410,1,0,0,2,52,66036,65656,65615]
e: 66410
d: [2,66411,2,2,1,1,40,65656,66110]
d: [2,66415,2,0,0,2,21,66411,66122]
e: 66415
d: [2,66412,1,2,1,1,40,65656,65615]
d: [2,66417,1,0,0,2,21,66412,66036]
e: 66417
r: 65629
r: 65656
r: 66011
a: [3,65628,0,2,1,65609]
s: 65628
a: [3,65625,0,4,1,65609]
s: 65625
a: [3,66010,0,7,1,55]
s: 66010
d: [2,66420,1,0,0,2,52,65968,66010,65620]
e: 66420
d: [2,66421,1,2,1,1,40,66010,65620]
d: [2,66425,1,0,0,2,21,66421,65968]
e: 66425
r: 65628
r: 65637
r: 65999
r: 66010
s: 65629
a: [3,65636,0,2,1,65609]
s: 65636
a: [3,65998,0,6,1,59]
s: 65998
d: [2,66427,1,0,0,2,52,65968,65998,65620]
e: 66427
d: [2,66428,1,0,0,2,52,65958,65998,65619]
e: 66428
d: [2,66429,1,2,1,1,40,65998,65620]
d: [2,66434,1,0,0,2,21,66429,65968]
e: 66434
d: [2,66430,1,2,1,1,40,65998,65619]
d: [2,66436,1,0,0,2,21,66430,65958]
e: 66436
r: 65625
r: 65636
r: 65998
s: 65637
a: [3,65624,0,4,1,65609]
s: 65624
d: [2,66438,2,0,0,2,52,66122,65624,66110]
e: 66438
d: [2,66439,1,0,0,2,52,66096,65624,65615]
e: 66439
d: [2,66440,2,2,1,1,40,65624,66110]
d: [2,66444,2,0,0,2,21,66440,66122]
e: 66444
d: [2,66441,1,2,1,1,40,65624,65615]
d: [2,66446,1,0,0,2,21,66441,66096]
e: 66446
r: 65624
r: 65633
r: 66083
r: 66086
r: 66105
r: 66139
r: 66146
r: 66172
r: 66022
r: 66027
s: 66026
s: 66087
s: 65632
d: [2,66452,2,0,0,2,52,66122,65632,66110]
e: 66452
d: [2,66453,2,2,1,1,40,65632,66110]
d: [2,66456,2,0,0,2,21,66453,66122]
e: 66456
r: 65632
r: 65668
r: 66179
s: 65669
a: [3,65649,0,4,1,65609]
s: 65649
s: 66023
a: [3,66104,1,2,1,66101]
s: 66104
s: 65999
a: [3,66145,1,2,1,66131]
s: 66145
s: 66011
s: 66082
s: 66140
a: [3,66365,0,7,1,66362]
s: 66365
s: 66173
s: 66017
a: [3,66178,1,3,1,66153]
s: 66178
d: [2,66477,2,0,0,2,52,65958,66178,65611]
e: 66477
d: [2,66478,2,2,1,1,40,66178,65611]
d: [2,66480,2,0,0,2,21,66478,65958]
e: 66480
r: 65955
r: 65958
r: 65677
r: 66005
s: 65676
s: 65959
a: [3,66374,2,3,1,66370]
s: 66374
s: 65954
d: [2,66487,2,0,0,2,52,66122,65954,66110]
e: 66487
d: [2,66488,2,2,1,1,40,65954,66110]
d: [2,66491,2,0,0,2,21,66488,66122]
e: 66491
r: 66122
r: 66082
r: 66087
r: 66140
r: 66374
s: 66086
s: 66123
s: 65633
s: 66139
d: [2,66499,2,0,0,2,52,66145,66139,65611]
e: 66499
d: [2,66500,2,2,1,1,40,66139,65611]
d: [2,66502,2,0,0,2,21,66500,66145]
e: 66502
r: 66145
r: 66173
s: 66004
s: 66083
s: 66146
s: 66172
d: [2,66509,2,0,0,2,52,66104,66172,65611]
e: 66509
d: [2,66510,2,2,1,1,40,66172,65611]
d: [2,66512,2,0,0,2,21,66510,66104]
e: 66512
r: 66023
r: 66026
r: 66104
r: 65649
r: 65676
r: 65945
r: 65948
r: 66004
r: 66135
r: 66043
r: 66046
r: 66172
s: 65676
s: 65948
s: 66026
a: [3,66047,0,2,1,47]
s: 66047
s: 65649
s: 66104
s: 65945
s: 66023
s: 66135
a: [3,66042,0,6,1,47]
s: 66042
s: 66173
s: 66004
d: [2,66371,1,4,1,1,40,66366,65614]
a: [3,66379,1,4,1,66371]
s: 66379
d: [2,66534,2,0,0,2,52,66056,66379,65615]
e: 66534
d: [2,66536,2,2,1,1,40,66379,65615]
d: [2,66538,2,0,0,2,21,66536,66056]
e: 66538
r: 66365
r: 66379
s: 66375
a: [3,65673,0,4,1,65609]
s: 65673
d: [2,66543,2,0,0,3,52,66047,66123,65673,65616]
e: 66543
d: [2,66544,1,4,2,1,40,65673,65616]
d: [2,66546,1,2,1,2,21,66544,66123]
d: [2,66547,1,0,0,3,21,66546,66047]
e: 66547
r: 65673
r: 66033
r: 66036
a: [3,66037,0,2,1,49]
s: 66037
s: 65656
a: [3,66032,0,6,1,49]
s: 66032
d: [2,66551,2,0,0,2,52,66202,66032,65620]
e: 66551
d: [2,66552,2,0,0,2,52,66104,66032,65619]
e: 66552
d: [2,66553,1,3,1,1,40,66032,65620]
d: [2,66558,1,0,0,2,21,66553,66202]
e: 66558
d: [2,66554,1,2,1,1,40,66032,65619]
d: [2,66560,1,0,0,2,21,66554,66104]
e: 66560
r: 65645
r: 65656
r: 65975
r: 66032
r: 66037
s: 65644
s: 66036
s: 65641
s: 65974
d: [2,66563,1,0,0,2,52,65968,65974,65620]
e: 66563
d: [2,66564,1,0,0,2,52,65948,65974,65618]
e: 66564
d: [2,66565,1,2,1,1,40,65974,65620]
d: [2,66570,1,0,0,2,21,66565,65968]
e: 66570
d: [2,66567,1,2,1,1,40,65974,65618]
d: [2,66572,1,0,0,2,21,66567,65948]
e: 66572
r: 65641
r: 65644
r: 65974
r: 66063
r: 66066
s: 65645
a: [3,66067,0,2,1,43]
s: 66067
s: 65640
a: [3,66380,1,4,1,66371]
s: 66380
s: 65975
s: 66033
a: [3,66062,0,6,1,43]
s: 66062
d: [2,66579,2,0,0,2,52,66104,66062,65619]
e: 66579
d: [2,66580,2,0,0,2,52,66139,66062,65618]
e: 66580
d: [2,66582,1,2,1,1,40,66062,65619]
d: [2,66586,1,0,0,2,21,66582,66104]
e: 66586
d: [2,66583,1,3,1,1,40,66062,65618]
d: [2,66588,1,0,0,2,21,66583,66139]
e: 66588
r: 65640
r: 65653
r: 65981
r: 66062
r: 66067
s: 65652
s: 66066
s: 65980
d: [2,66590,1,0,0,2,52,65948,65980,65618]
e: 66590
d: [2,66593,1,2,1,1,40,65980,65618]
d: [2,66595,1,0,0,2,21,66593,65948]
e: 66595
r: 65652
r: 65669
r: 65980
r: 66017
s: 65653
s: 65668
s: 65981
s: 66063
s: 66366
a: [3,66016,0,8,1,53]
s: 66016
d: [2,66601,2,0,0,2,52,66178,66016,65619]
e: 66601
d: [2,66603,1,3,1,1,40,66016,65619]
d: [2,66606,1,0,0,2,21,66603,66178]
e: 66606
r: 65649
r: 65653
r: 65661
r: 65668
r: 65676
r: 65945
r: 65948
r: 65981
r: 65993
r: 66004
r: 66016
r: 66023
r: 66026
r: 66042
r: 66047
r: 66104
r: 66135
r: 66173
s: 65652
s: 65660
s: 65669
s: 65677
s: 65949
s: 66027
s: 66046
s: 66105
s: 65648
s: 65657
s: 66134
s: 65944
s: 66172
s: 66022
s: 65980
s: 66043
s: 65992
d: [2,66622,2,0,0,2,52,66178,65992,65619]
e: 66622
d: [2,66624,1,3,1,1,40,65992,65619]
d: [2,66627,1,0,0,2,21,66624,66178]
e: 66627
r: 65652
r: 65660
r: 65968
r: 65980
r: 65992
r: 66115
r: 66197
r: 65629
r: 65648
r: 65954
r: 65959
r: 66011
r: 66022
r: 66027
r: 66178
r: 66185
r: 65645
r: 65975
r: 66063
r: 66066
r: 66139
r: 66146
s: 65628
s: 65644
s: 65653
s: 65661
s: 65959
s: 65969
s: 66026
s: 66067
s: 66145
s: 65625
s: 66185
s: 65649
s: 65954
s: 66140
s: 65974
s: 66178
s: 66023
s: 66114
s: 66196
s: 65981
s: 66062
s: 65993
s: 66005
s: 66010
d: [2,66658,2,0,0,2,52,66178,66010,65619]
e: 66658
d: [2,66659,2,0,0,2,52,66134,66010,65618]
e: 66659
d: [2,66661,1,3,1,1,40,66010,65619]
d: [2,66665,1,0,0,2,21,66661,66178]
e: 66665
d: [2,66662,1,3,1,1,40,66010,65618]
d: [2,66667,1,0,0,2,21,66662,66134]
e: 66667
r: 65628
r: 66010
s: 65629
s: 65641
d: [2,66669,2,0,0,3,52,66067,66123,65641,65616]
e: 66669
d: [2,66670,1,4,2,1,40,65641,65616]
d: [2,66672,1,2,1,2,21,66670,66123]
d: [2,66673,1,0,0,3,21,66672,66067]
e: 66673
r: 65637
r: 65641
r: 65999
s: 65636
s: 65998
d: [2,66675,2,0,0,2,52,66134,65998,65618]
e: 66675
d: [2,66678,1,3,1,1,40,65998,65618]
d: [2,66680,1,0,0,2,21,66678,66134]
e: 66680
r: 65636
r: 65649
r: 65998
r: 66023
r: 66026
r: 66062
r: 66067
r: 66140
r: 66145
r: 65625
r: 65644
r: 65944
r: 65949
r: 65954
r: 65959
r: 65974
r: 66093
r: 66096
r: 66105
r: 66134
r: 66172
r: 66178
r: 66185
r: 66202
s: 65637
s: 65645
s: 65948
s: 65958
s: 66026
s: 66066
s: 66097
s: 66104
s: 65624
s: 65649
s: 66145
s: 65945
s: 66184
s: 65955
s: 65975
s: 66135
s: 66023
s: 66140
s: 65999
s: 66063
s: 66173
s: 66011
r: 65624
r: 65629
r: 66011
s: 65628
s: 65625
d: [2,66706,2,0,0,3,52,66097,66123,65625,65616]
e: 66706
d: [2,66707,1,4,2,1,40,65625,65616]
d: [2,66709,1,2,1,2,21,66707,66123]
d: [2,66710,1,0,0,3,21,66709,66097]
e: 66710
r: 66097
s: 66096
s: 66179
s: 66010
s: 66203
s: 66017
s: 66093
d: [2,66719,2,0,0,4,52,66140,66173,66203,66093,65621]
e: 66719
d: [2,66720,1,9,3,1,40,66093,65621]
d: [2,66722,1,6,2,2,21,66720,66203]
d: [2,66723,1,3,1,3,21,66722,66173]
d: [2,66724,1,0,0,4,21,66723,66140]
e: 66724
r: 65625
r: 66073
r: 66076
r: 66093
r: 66096
s: 66077
s: 66097
s: 65624
s: 65680
s: 66072
d: [2,66730,2,0,0,2,52,66104,66072,65619]
e: 66730
d: [2,66732,1,2,1,1,40,66072,65619]
d: [2,66735,1,0,0,2,21,66732,66104]
e: 66735
r: 65680
r: 66043
r: 66046
r: 66072
r: 66077
s: 66047
s: 66076
s: 65672
s: 66042
d: [2,66739,2,0,0,2,52,66184,66042,65620]
e: 66739
d: [2,66740,2,0,0,2,52,66145,66042,65618]
e: 66740
d: [2,66741,1,2,1,1,40,66042,65620]
d: [2,66746,1,0,0,2,21,66741,66184]
e: 66746
d: [2,66743,1,2,1,1,40,66042,65618]
d: [2,66748,1,0,0,2,21,66743,66145]
e: 66748
r: 65672
r: 65685
r: 65987
r: 66042
r: 66047
s: 65684
s: 66046
a: [3,65681,0,4,1,65609]
s: 65681
s: 65986
d: [2,66751,1,0,0,2,52,65948,65986,65618]
e: 66751
d: [2,66754,1,2,1,1,40,65986,65618]
d: [2,66756,1,0,0,2,21,66754,65948]
e: 66756
r: 65669
r: 65684
r: 65986
r: 66017
s: 65668
s: 65685
s: 65987
s: 66043
s: 66073
s: 66016
d: [2,66761,2,0,0,2,52,66196,66016,65620]
e: 66761
d: [2,66762,1,3,1,1,40,66016,65620]
d: [2,66766,1,0,0,2,21,66762,66196]
e: 66766
r: 65624
r: 65628
r: 65969
r: 66010
r: 66053
r: 66056
r: 66097
r: 66114
r: 66140
r: 66145
r: 66196
r: 65661
r: 65945
r: 65948
r: 65993
r: 66033
r: 66036
r: 66104
r: 66135
r: 66173
r: 65668
r: 65685
r: 65987
r: 66016
r: 66073
r: 66076
r: 65657
r: 65681
s: 65629
s: 65661
s: 65669
s: 65684
s: 65949
s: 65969
s: 66037
s: 66056
s: 66077
s: 66096
s: 65656
s: 66105
s: 65680
s: 66145
s: 65944
s: 66114
s: 66134
s: 65986
s: 66140
s: 66032
d: [2,66798,2,0,0,2,52,66145,66032,65618]
e: 66798
d: [2,66801,1,2,1,1,40,66032,65618]
d: [2,66803,1,0,0,2,21,66801,66145]
e: 66803
r: 65656
r: 66032
r: 66037
r: 66063
r: 66066
s: 66036
s: 66067
s: 65640
s: 66172
s: 66033
s: 66062
d: [2,66809,2,0,0,2,52,66184,66062,65620]
e: 66809
d: [2,66810,1,2,1,1,40,66062,65620]
d: [2,66814,1,0,0,2,21,66810,66184]
e: 66814
r: 65637
r: 65640
r: 65680
r: 65684
r: 65944
r: 65949
r: 65955
r: 65958
r: 65986
r: 65999
r: 66062
r: 66067
r: 66077
r: 66083
r: 66086
r: 66105
r: 66134
r: 66172
r: 66179
r: 66184
r: 66203
r: 65633
s: 65637
s: 65685
s: 65948
s: 65959
s: 66066
s: 66076
s: 66087
s: 66104
s: 65632
s: 65945
s: 66185
s: 65954
s: 66135
s: 65987
s: 65999
d: [2,66829,2,0,0,4,52,65969,65959,66135,65999,65621]
e: 66829
d: [2,66830,1,7,3,1,40,65999,65621]
d: [2,66832,1,5,2,2,21,66830,65969]
d: [2,66833,1,3,1,3,21,66832,65959]
d: [2,66834,1,0,0,4,21,66833,66135]
e: 66834
i: [1,49226,0,51,16,1,0,0] t4_gate_3
i: [1,12362,0,51,16,1,0,0] t4_gate_3
i: [1,3146,0,51,16,1,0,0] t4_gate_3
i: [1,842,0,51,16,1,0,0] t4_gate_3
i: [1,266,0,51,16,1,0,0] t4_gate_3
i: [1,122,0,51,16,1,0,0] t4_gate_3
i: [1,86,0,51,16,1,0,0] t4_gate_3
i: [1,77,0,51,16,1,0,0] t4_gate_3
i: [1,72,0,6,2,1,0,0] t4_gate_3
i: [1,70,0,6,2,1,0,0] t4_gate_3
i: [1,68,0,7,2,1,0,0] t4_gate_3
i: [1,66,0,7,2,1,0,0] t4_gate_3
i: [1,64,0,8,2,1,0,0] t4_gate_3
i: [1,62,0,8,2,1,0,0] t4_gate_3
i: [1,60,0,9,2,1,0,0] t4_gate_3
i: [1,58,0,8,2,1,0,0] t4_gate_3
i: [1,56,0,9,2,1,0,0] t4_gate_3
i: [1,54,0,9,2,1,0,0] t4_gate_3
i: [1,52,0,10,2,1,0,0] t4_gate_3
i: [1,50,0,7,2,1,0,0] t4_gate_3
i: [1,48,0,8,2,1,0,0] t4_gate_3
i: [1,46,0,8,2,1,0,0] t4_gate_3
i: [1,44,0,9,2,1,0,0] t4_gate_3
i: [1,42,0,8,2,1,0,0] t4_gate_3
i: [1,40,0,9,2,1,0,0] t4_gate_3
i: [1,38,0,9,2,1,0,0] t4_gate_3
i: [1,36,0,10,2,1,0,0] t4_gate_3
d: [2,66130,1,3,1,1,40,65618,66092]
d: [2,66156,1,3,1,1,40,65619,66092]
d: [2,66190,1,3,1,1,40,65620,66004]
d: [2,66192,1,3,1,1,40,65620,66092]
d: [2,66193,1,3,1,1,40,65620,66052]
d: [2,66189,1,2,1,1,40,65620,65986]
d: [2,66261,2,2,1,1,40,66203,65610]
d: [2,66282,1,2,1,1,40,66022,65620]
d: [2,66305,1,8,3,1,40,65621,65993]
d: [2,66319,1,5,2,2,21,66305,66179]
d: [2,66320,1,3,1,3,21,66319,65949]
d: [2,66329,1,3,1,1,40,65992,65620]
d: [2,66333,2,2,1,1,40,66134,65611]
d: [2,66331,1,2,1,1,40,65992,65618]
a: [3,65664,0,7,1,65609]
d: [2,66387,0,7,1,1,22,65664,65612]
d: [2,66296,1,7,3,1,40,65621,66063]
d: [2,66309,1,5,2,2,21,66296,66185]
d: [2,66283,1,2,1,1,40,66022,65619]
d: [2,66284,1,2,1,1,40,66022,65618]
d: [2,66459,1,6,3,1,40,66023,65621]
d: [2,66460,1,4,2,2,21,66459,66185]
d: [2,66465,1,3,1,1,40,66082,65619]
d: [2,66466,1,3,1,1,40,66082,65618]
d: [2,66470,1,7,2,1,40,66365,65617]
d: [2,66473,1,9,3,1,40,66017,65621]
d: [2,66474,1,6,2,2,21,66473,66197]
d: [2,66475,1,3,1,3,21,66474,66135]
d: [2,66486,3,2,1,1,40,66374,65610]
d: [2,66286,1,5,2,1,40,66115,65616]
d: [2,66458,1,4,2,1,40,65649,65616]
d: [2,66514,1,2,1,2,21,66458,66123]
d: [2,66368,1,4,2,1,40,65945,65616]
d: [2,66517,1,2,1,2,21,66368,66123]
d: [2,66527,1,3,1,1,40,66042,65619]
d: [2,66540,0,7,1,1,22,65665,65612]
d: [2,66287,2,2,1,1,40,66202,65611]
d: [2,66508,1,8,3,1,40,66083,65621]
d: [2,66621,1,7,3,1,40,66043,65621]
d: [2,66578,1,7,3,1,40,66033,65621]
d: [2,66304,1,8,3,1,40,65621,66053]
d: [2,66702,1,8,3,1,40,66011,65621]
d: [2,66703,1,6,2,2,21,66702,65969]
d: [2,66704,1,3,1,3,21,66703,66135]
d: [2,66576,2,4,2,1,40,66380,65616]
d: [2,66768,2,2,1,2,21,66576,66123]
d: [2,66771,0,5,1,1,22,65965,65613]
d: [2,66758,1,7,3,1,40,65987,65621]
d: [2,66773,1,4,2,2,21,66758,66179]
d: [2,66610,1,4,2,1,40,65657,65616]
d: [2,66775,1,2,1,2,21,66610,66123]
d: [2,66760,1,8,3,1,40,66073,65621]
d: [2,66777,1,5,2,2,21,66760,66203]
d: [2,66764,1,3,1,1,40,66016,65618]
d: [2,66750,1,4,2,1,40,65681,65616]
d: [2,66780,1,2,1,2,21,66750,66123]
a: [3,65964,0,5,1,69]
d: [2,66782,0,5,1,1,22,65964,65613]
d: [2,66790,2,3,1,1,40,66114,66110]
d: [2,66652,1,7,3,1,40,65981,65621]
d: [2,66657,1,8,3,1,40,66005,65621]
d: [2,66695,1,6,3,1,40,65975,65621]
d: [2,66818,1,4,2,2,21,66695,65969]
d: [2,66498,1,4,2,1,40,65633,65616]
d: [2,66820,1,2,1,2,21,66498,66123]
f: 66834,66820,66818,66657,66652,66814,66803,66790,66782,66780,66764,66777,66775,66773,66771,66762,66768,66756,66748,66746,66735,66724,66709,66704,66304,66578,66621,66508,66680,66673,66667,66665,66287,66627,66606,66595,66588,66586,66572,66570,66560,66558,66547,66540,66538,66527,66517,66510,66514,66500,66286,66486,66488,66478,66475,66470,66466,66465,66460,66456,66284,66283,66309,66446,66444,66436,66434,66425,66417,66415,66407,66405,66397,66395,66387,66385,66371,66370,66360,66358,66331,66333,66346,66344,66329,66327,66320,66282,66277,66261,66263,66259,66242,66189,66236,66230,66227,66222,66217,66193,66192,66191,66190,66188,66181,66153,66156,66157,66151,66152,66154,66131,66130,66129,66128,66119,66101,36,37,38,39,40,41,42,43,44,45,46,47,48,49,50,51,52,53,54,55,56,57,58,59,60,61,62,63,64,65,66,67,68,69,70,71,72,73,77,86,122,266,842,3146,12362,49226
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_3__t4_gate_3
% SZS output start Proof for gate_3__t4_gate_3
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ! [X22] : ~(~(~(v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X2,X22))) & ~(v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X5,X22))) & ~(v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22)))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X6,X22))) & ~(v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X7,X22))) & ~(v1_xboole_0(k2_gate_1(X7,X22)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X3,X22))) & ~(v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X1,X22))) & ~(v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) & ~(v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k2_gate_1(X17,X22))) & ~(v1_xboole_0(k2_gate_1(X17,X22)) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k2_gate_1(X16,X22))) & ~(v1_xboole_0(k2_gate_1(X16,X22)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [input t4_gate_3]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ! [X22] : ~(~(~(v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X2,X22))) & ~(v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X5,X22))) & ~(v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22)))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X6,X22))) & ~(v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X7,X22))) & ~(v1_xboole_0(k2_gate_1(X7,X22)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X3,X22))) & ~(v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X1,X22))) & ~(v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) & ~(v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k2_gate_1(X17,X22))) & ~(v1_xboole_0(k2_gate_1(X17,X22)) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k2_gate_1(X16,X22))) & ~(v1_xboole_0(k2_gate_1(X16,X22)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [negated conjecture 1]
3. ! [X0,X1] : k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [input commutativity_k2_gate_1]
4. ! [X0,X1] : k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [input commutativity_k3_gate_1]
9. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc17_gate_1]
10. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc18_gate_1]
11. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X0,X2))) [input fc19_gate_1]
12. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
13. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X2,X0))) [input fc20_gate_1]
14. ! [X0] : (~v1_xboole_0(X0) => v1_xboole_0(k1_gate_1(X0))) [input fc2_gate_1]
15. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_gate_1(X0,X1))) [input fc3_gate_1]
16. ! [X0,X1] : (v1_xboole_0(X0) => v1_xboole_0(k2_gate_1(X0,X1))) [input fc4_gate_1]
17. ! [X0,X1] : (~v1_xboole_0(X1) => ~v1_xboole_0(k3_gate_1(X0,X1))) [input fc5_gate_1]
18. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k3_gate_1(X0,X1))) [input fc6_gate_1]
19. ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22] : ~(~(~(v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X2,X22))) & ~(v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X5,X22))) & ~(v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22)))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X6,X22))) & ~(v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X7,X22))) & ~(v1_xboole_0(k2_gate_1(X7,X22)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X3,X22))) & ~(v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X1,X22))) & ~(v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) & ~(v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k2_gate_1(X17,X22))) & ~(v1_xboole_0(k2_gate_1(X17,X22)) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k2_gate_1(X16,X22))) & ~(v1_xboole_0(k2_gate_1(X16,X22)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [flattening 2]
20. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22] : (((v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X2,X22))) | (v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X13)) | (v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X5,X22))) | (v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X10)) | (v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22)))) | (v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22))) & ~v1_xboole_0(X8)) | (v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X6,X22))) | (v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X12)) | (v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X7,X22))) | (v1_xboole_0(k2_gate_1(X7,X22)) & ~v1_xboole_0(X14)) | (v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X3,X22))) | (v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X15)) | (v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X1,X22))) | (v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X11)) | (v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) | (v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & (~v1_xboole_0(X21) | v1_xboole_0(k2_gate_1(X17,X22))) & (~v1_xboole_0(k2_gate_1(X17,X22)) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(k2_gate_1(X16,X22))) & (~v1_xboole_0(k2_gate_1(X16,X22)) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22)) | v1_xboole_0(X19)) & (~v1_xboole_0(X15) | v1_xboole_0(k8_gate_1(X21,X20,X19))) & (~v1_xboole_0(k8_gate_1(X21,X20,X19)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) | v1_xboole_0(X10)) & (~v1_xboole_0(X9) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) | v1_xboole_0(X9)) & (~v1_xboole_0(X8) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X8)) & (~v1_xboole_0(X7) | v1_xboole_0(k8_gate_1(X18,X17,X16))) & (~v1_xboole_0(k8_gate_1(X18,X17,X16)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X0))) [ennf transformation 19]
21. ! [X0] : (v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0)) [ennf transformation 14]
22. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 12]
23. ! [X0,X1] : (~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1)) [ennf transformation 17]
24. ! [X0,X1] : (v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 16]
25. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 15]
26. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 25]
27. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 18]
28. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 27]
29. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0)) [ennf transformation 10]
30. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0)) [ennf transformation 11]
31. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0)) [ennf transformation 13]
32. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 9]
33. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 32]
34. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22] : (((v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X2,X22))) | (v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X13)) | (v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X5,X22))) | (v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X10)) | (v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22)))) | (v1_xboole_0(k3_gate_1(k2_gate_1(X4,X22),k1_gate_1(X22))) & ~v1_xboole_0(X8)) | (v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X6,X22))) | (v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X12)) | (v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X7,X22))) | (v1_xboole_0(k2_gate_1(X7,X22)) & ~v1_xboole_0(X14)) | (v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X3,X22))) | (v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X15)) | (v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X1,X22))) | (v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X11)) | (v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) | (v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & (~v1_xboole_0(X21) | v1_xboole_0(k2_gate_1(X17,X22))) & (~v1_xboole_0(k2_gate_1(X17,X22)) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(k2_gate_1(X16,X22))) & (~v1_xboole_0(k2_gate_1(X16,X22)) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X18),X22)) | v1_xboole_0(X19)) & (~v1_xboole_0(X15) | v1_xboole_0(k8_gate_1(X21,X20,X19))) & (~v1_xboole_0(k8_gate_1(X21,X20,X19)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) | v1_xboole_0(X10)) & (~v1_xboole_0(X9) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) | v1_xboole_0(X9)) & (~v1_xboole_0(X8) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X8)) & (~v1_xboole_0(X7) | v1_xboole_0(k8_gate_1(X18,X17,X16))) & (~v1_xboole_0(k8_gate_1(X18,X17,X16)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X0))) => (((v1_xboole_0(sK13) & ~v1_xboole_0(k2_gate_1(sK2,sK22))) | (v1_xboole_0(k2_gate_1(sK2,sK22)) & ~v1_xboole_0(sK13)) | (v1_xboole_0(sK10) & ~v1_xboole_0(k2_gate_1(sK5,sK22))) | (v1_xboole_0(k2_gate_1(sK5,sK22)) & ~v1_xboole_0(sK10)) | (v1_xboole_0(sK8) & ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22)))) | (v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) & ~v1_xboole_0(sK8)) | (v1_xboole_0(sK12) & ~v1_xboole_0(k2_gate_1(sK6,sK22))) | (v1_xboole_0(k2_gate_1(sK6,sK22)) & ~v1_xboole_0(sK12)) | (v1_xboole_0(sK14) & ~v1_xboole_0(k2_gate_1(sK7,sK22))) | (v1_xboole_0(k2_gate_1(sK7,sK22)) & ~v1_xboole_0(sK14)) | (v1_xboole_0(sK15) & ~v1_xboole_0(k2_gate_1(sK3,sK22))) | (v1_xboole_0(k2_gate_1(sK3,sK22)) & ~v1_xboole_0(sK15)) | (v1_xboole_0(sK11) & ~v1_xboole_0(k2_gate_1(sK1,sK22))) | (v1_xboole_0(k2_gate_1(sK1,sK22)) & ~v1_xboole_0(sK11)) | (v1_xboole_0(sK9) & ~v1_xboole_0(k2_gate_1(sK0,sK22))) | (v1_xboole_0(k2_gate_1(sK0,sK22)) & ~v1_xboole_0(sK9))) & (~v1_xboole_0(sK21) | v1_xboole_0(k2_gate_1(sK17,sK22))) & (~v1_xboole_0(k2_gate_1(sK17,sK22)) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(k2_gate_1(sK16,sK22))) & (~v1_xboole_0(k2_gate_1(sK16,sK22)) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k2_gate_1(k1_gate_1(sK18),sK22))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK18),sK22)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0))) [choice axiom]
35. ((v1_xboole_0(sK13) & ~v1_xboole_0(k2_gate_1(sK2,sK22))) | (v1_xboole_0(k2_gate_1(sK2,sK22)) & ~v1_xboole_0(sK13)) | (v1_xboole_0(sK10) & ~v1_xboole_0(k2_gate_1(sK5,sK22))) | (v1_xboole_0(k2_gate_1(sK5,sK22)) & ~v1_xboole_0(sK10)) | (v1_xboole_0(sK8) & ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22)))) | (v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) & ~v1_xboole_0(sK8)) | (v1_xboole_0(sK12) & ~v1_xboole_0(k2_gate_1(sK6,sK22))) | (v1_xboole_0(k2_gate_1(sK6,sK22)) & ~v1_xboole_0(sK12)) | (v1_xboole_0(sK14) & ~v1_xboole_0(k2_gate_1(sK7,sK22))) | (v1_xboole_0(k2_gate_1(sK7,sK22)) & ~v1_xboole_0(sK14)) | (v1_xboole_0(sK15) & ~v1_xboole_0(k2_gate_1(sK3,sK22))) | (v1_xboole_0(k2_gate_1(sK3,sK22)) & ~v1_xboole_0(sK15)) | (v1_xboole_0(sK11) & ~v1_xboole_0(k2_gate_1(sK1,sK22))) | (v1_xboole_0(k2_gate_1(sK1,sK22)) & ~v1_xboole_0(sK11)) | (v1_xboole_0(sK9) & ~v1_xboole_0(k2_gate_1(sK0,sK22))) | (v1_xboole_0(k2_gate_1(sK0,sK22)) & ~v1_xboole_0(sK9))) & (~v1_xboole_0(sK21) | v1_xboole_0(k2_gate_1(sK17,sK22))) & (~v1_xboole_0(k2_gate_1(sK17,sK22)) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(k2_gate_1(sK16,sK22))) & (~v1_xboole_0(k2_gate_1(sK16,sK22)) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k2_gate_1(k1_gate_1(sK18),sK22))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK18),sK22)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0)) [skolemisation 20,34]
36. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0) [cnf transformation 35] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. ~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) [cnf transformation 35] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. ~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
41. ~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. ~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
44. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. ~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. ~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. ~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. ~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. ~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7) [cnf transformation 35] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. ~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) [cnf transformation 35] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8) [cnf transformation 35] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. ~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) [cnf transformation 35] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. ~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
56. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
57. ~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. ~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
60. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
61. ~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
62. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
63. ~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
64. ~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. ~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) [cnf transformation 35] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
66. ~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15) [cnf transformation 35] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
67. ~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) [cnf transformation 35] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
68. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK18),sK22)) | v1_xboole_0(sK19) [cnf transformation 35] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. ~v1_xboole_0(sK19) | v1_xboole_0(k2_gate_1(k1_gate_1(sK18),sK22)) [cnf transformation 35] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. ~v1_xboole_0(k2_gate_1(sK16,sK22)) | v1_xboole_0(sK20) [cnf transformation 35] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
71. ~v1_xboole_0(sK20) | v1_xboole_0(k2_gate_1(sK16,sK22)) [cnf transformation 35] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
72. ~v1_xboole_0(k2_gate_1(sK17,sK22)) | v1_xboole_0(sK21) [cnf transformation 35] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
73. ~v1_xboole_0(sK21) | v1_xboole_0(k2_gate_1(sK17,sK22)) [cnf transformation 35] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
77. ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK7,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK11) | v1_xboole_0(sK9) | v1_xboole_0(k2_gate_1(sK0,sK22)) [cnf transformation 35] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
86. ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK7,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK15) | v1_xboole_0(sK11) | v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 35] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
122. ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK7,sK22)) | ~v1_xboole_0(sK14) | v1_xboole_0(sK15) | v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 35] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
266. ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK12) | v1_xboole_0(sK14) | v1_xboole_0(k2_gate_1(sK7,sK22)) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 35] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
842. ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | v1_xboole_0(sK12) | v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(k2_gate_1(sK7,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 35] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
3146. ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK10) | v1_xboole_0(sK8) | v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK7,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 35] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
12362. ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK13) | v1_xboole_0(sK10) | v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK7,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 35] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49226. v1_xboole_0(sK13) | v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK7,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 35] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65610. v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0) [cnf transformation 21] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65611. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 22] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65612. k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
65613. k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
65614. ~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1) [cnf transformation 23] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65615. v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 24] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65616. ~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65617. v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 28] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65618. v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0) [cnf transformation 29] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65619. v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0) [cnf transformation 30] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65620. v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0) [cnf transformation 31] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65621. ~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 33] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65623. 1 <=> v1_xboole_0(k2_gate_1(sK0,sK22)) [avatar definition]
65624. ~v1_xboole_0(k2_gate_1(sK0,sK22)) <- (~1) [avatar component clause 65623] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65625. v1_xboole_0(k2_gate_1(sK0,sK22)) <- (1) [avatar component clause 65623] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65627. 2 <=> v1_xboole_0(sK9) [avatar definition]
65631. 3 <=> v1_xboole_0(k2_gate_1(sK1,sK22)) [avatar definition]
65632. ~v1_xboole_0(k2_gate_1(sK1,sK22)) <- (~3) [avatar component clause 65631] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65633. v1_xboole_0(k2_gate_1(sK1,sK22)) <- (3) [avatar component clause 65631] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65635. 4 <=> v1_xboole_0(sK11) [avatar definition]
65639. 5 <=> v1_xboole_0(k2_gate_1(sK3,sK22)) [avatar definition]
65640. ~v1_xboole_0(k2_gate_1(sK3,sK22)) <- (~5) [avatar component clause 65639] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65641. v1_xboole_0(k2_gate_1(sK3,sK22)) <- (5) [avatar component clause 65639] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65643. 6 <=> v1_xboole_0(sK15) [avatar definition]
65647. 7 <=> v1_xboole_0(k2_gate_1(sK7,sK22)) [avatar definition]
65648. ~v1_xboole_0(k2_gate_1(sK7,sK22)) <- (~7) [avatar component clause 65647] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65649. v1_xboole_0(k2_gate_1(sK7,sK22)) <- (7) [avatar component clause 65647] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65651. 8 <=> v1_xboole_0(sK14) [avatar definition]
65655. 9 <=> v1_xboole_0(k2_gate_1(sK6,sK22)) [avatar definition]
65656. ~v1_xboole_0(k2_gate_1(sK6,sK22)) <- (~9) [avatar component clause 65655] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65657. v1_xboole_0(k2_gate_1(sK6,sK22)) <- (9) [avatar component clause 65655] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65659. 10 <=> v1_xboole_0(sK12) [avatar definition]
65663. 11 <=> v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) [avatar definition]
65664. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) <- (~11) [avatar component clause 65663] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65665. v1_xboole_0(k3_gate_1(k2_gate_1(sK4,sK22),k1_gate_1(sK22))) <- (11) [avatar component clause 65663] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65667. 12 <=> v1_xboole_0(sK8) [avatar definition]
65671. 13 <=> v1_xboole_0(k2_gate_1(sK5,sK22)) [avatar definition]
65672. ~v1_xboole_0(k2_gate_1(sK5,sK22)) <- (~13) [avatar component clause 65671] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65673. v1_xboole_0(k2_gate_1(sK5,sK22)) <- (13) [avatar component clause 65671] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65675. 14 <=> v1_xboole_0(sK10) [avatar definition]
65679. 15 <=> v1_xboole_0(k2_gate_1(sK2,sK22)) [avatar definition]
65680. ~v1_xboole_0(k2_gate_1(sK2,sK22)) <- (~15) [avatar component clause 65679] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65681. v1_xboole_0(k2_gate_1(sK2,sK22)) <- (15) [avatar component clause 65679] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65683. 16 <=> v1_xboole_0(sK13) [avatar definition]
65813. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | 15 | 16 [avatar split clause 49226,65683,65679,65671,65675,65663,65667,65655,65659,65647,65651,65639,65643,65631,65635,65623,65627]
65877. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | 13 | 14 | ~16 | ~15 [avatar split clause 12362,65679,65683,65675,65671,65663,65667,65655,65659,65647,65651,65639,65643,65631,65635,65623,65627]
65909. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | 11 | 12 | ~14 | ~13 | ~16 | ~15 [avatar split clause 3146,65679,65683,65671,65675,65667,65663,65655,65659,65647,65651,65639,65643,65631,65635,65623,65627]
65925. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | 9 | 10 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 842,65679,65683,65671,65675,65663,65667,65659,65655,65647,65651,65639,65643,65631,65635,65623,65627]
65933. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | 7 | 8 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 266,65679,65683,65671,65675,65663,65667,65655,65659,65651,65647,65639,65643,65631,65635,65623,65627]
65937. ~2 | ~1 | ~4 | ~3 | 5 | 6 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 122,65679,65683,65671,65675,65663,65667,65655,65659,65647,65651,65643,65639,65631,65635,65623,65627]
65939. ~2 | ~1 | 3 | 4 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 86,65679,65683,65671,65675,65663,65667,65655,65659,65647,65651,65639,65643,65635,65631,65623,65627]
65940. 1 | 2 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 77,65679,65683,65671,65675,65663,65667,65655,65659,65647,65651,65639,65643,65631,65635,65627,65623]
65943. 17 <=> v1_xboole_0(k2_gate_1(sK17,sK22)) [avatar definition]
65944. ~v1_xboole_0(k2_gate_1(sK17,sK22)) <- (~17) [avatar component clause 65943] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65945. v1_xboole_0(k2_gate_1(sK17,sK22)) <- (17) [avatar component clause 65943] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65947. 18 <=> v1_xboole_0(sK21) [avatar definition]
65948. v1_xboole_0(sK21) <- (18) [avatar component clause 65947] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65949. ~v1_xboole_0(sK21) <- (~18) [avatar component clause 65947] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65950. 17 | ~18 [avatar split clause 73,65947,65943]
65951. 18 | ~17 [avatar split clause 72,65943,65947]
65953. 19 <=> v1_xboole_0(k2_gate_1(sK16,sK22)) [avatar definition]
65954. ~v1_xboole_0(k2_gate_1(sK16,sK22)) <- (~19) [avatar component clause 65953] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65955. v1_xboole_0(k2_gate_1(sK16,sK22)) <- (19) [avatar component clause 65953] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65957. 20 <=> v1_xboole_0(sK20) [avatar definition]
65958. v1_xboole_0(sK20) <- (20) [avatar component clause 65957] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65959. ~v1_xboole_0(sK20) <- (~20) [avatar component clause 65957] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65960. 19 | ~20 [avatar split clause 71,65957,65953]
65961. 20 | ~19 [avatar split clause 70,65953,65957]
65963. 21 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK18),sK22)) [avatar definition]
65964. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK18),sK22)) <- (~21) [avatar component clause 65963] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65965. v1_xboole_0(k2_gate_1(k1_gate_1(sK18),sK22)) <- (21) [avatar component clause 65963] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65967. 22 <=> v1_xboole_0(sK19) [avatar definition]
65968. v1_xboole_0(sK19) <- (22) [avatar component clause 65967] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65969. ~v1_xboole_0(sK19) <- (~22) [avatar component clause 65967] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65970. 21 | ~22 [avatar split clause 69,65967,65963]
65971. 22 | ~21 [avatar split clause 68,65963,65967]
65973. 23 <=> v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) [avatar definition]
65974. ~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) <- (~23) [avatar component clause 65973] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65975. v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) <- (23) [avatar component clause 65973] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65976. 23 | ~6 [avatar split clause 67,65643,65973]
65977. 6 | ~23 [avatar split clause 66,65973,65643]
65979. 24 <=> v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) [avatar definition]
65980. ~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) <- (~24) [avatar component clause 65979] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65981. v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) <- (24) [avatar component clause 65979] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65982. 24 | ~8 [avatar split clause 65,65651,65979]
65983. 8 | ~24 [avatar split clause 64,65979,65651]
65985. 25 <=> v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) [avatar definition]
65986. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) <- (~25) [avatar component clause 65985] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65987. v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) <- (25) [avatar component clause 65985] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65988. 25 | ~16 [avatar split clause 63,65683,65985]
65989. 16 | ~25 [avatar split clause 62,65985,65683]
65991. 26 <=> v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) [avatar definition]
65992. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) <- (~26) [avatar component clause 65991] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65993. v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) <- (26) [avatar component clause 65991] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65994. 26 | ~10 [avatar split clause 61,65659,65991]
65995. 10 | ~26 [avatar split clause 60,65991,65659]
65997. 27 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) [avatar definition]
65998. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) <- (~27) [avatar component clause 65997] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65999. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) <- (27) [avatar component clause 65997] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66000. 27 | ~4 [avatar split clause 59,65635,65997]
66001. 4 | ~27 [avatar split clause 58,65997,65635]
66003. 28 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) [avatar definition]
66004. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) <- (~28) [avatar component clause 66003] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66005. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) <- (28) [avatar component clause 66003] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66006. 28 | ~14 [avatar split clause 57,65675,66003]
66007. 14 | ~28 [avatar split clause 56,66003,65675]
66009. 29 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) [avatar definition]
66010. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) <- (~29) [avatar component clause 66009] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66011. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) <- (29) [avatar component clause 66009] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66012. 29 | ~2 [avatar split clause 55,65627,66009]
66013. 2 | ~29 [avatar split clause 54,66009,65627]
66015. 30 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) [avatar definition]
66016. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) <- (~30) [avatar component clause 66015] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66017. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) <- (30) [avatar component clause 66015] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66018. 30 | ~12 [avatar split clause 53,65667,66015]
66019. 12 | ~30 [avatar split clause 52,66015,65667]
66021. 31 <=> v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) [avatar definition]
66022. ~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) <- (~31) [avatar component clause 66021] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66023. v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) <- (31) [avatar component clause 66021] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66025. 32 <=> v1_xboole_0(sK7) [avatar definition]
66028. 31 | ~32 [avatar split clause 51,66025,66021]
66029. 32 | ~31 [avatar split clause 50,66021,66025]
66031. 33 <=> v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) [avatar definition]
66032. ~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) <- (~33) [avatar component clause 66031] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66033. v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) <- (33) [avatar component clause 66031] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66035. 34 <=> v1_xboole_0(sK6) [avatar definition]
66036. v1_xboole_0(sK6) <- (34) [avatar component clause 66035] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66038. 33 | ~34 [avatar split clause 49,66035,66031]
66039. 34 | ~33 [avatar split clause 48,66031,66035]
66041. 35 <=> v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) [avatar definition]
66042. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) <- (~35) [avatar component clause 66041] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66043. v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) <- (35) [avatar component clause 66041] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66045. 36 <=> v1_xboole_0(sK5) [avatar definition]
66046. v1_xboole_0(sK5) <- (36) [avatar component clause 66045] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66047. ~v1_xboole_0(sK5) <- (~36) [avatar component clause 66045] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66048. 35 | ~36 [avatar split clause 47,66045,66041]
66049. 36 | ~35 [avatar split clause 46,66041,66045]
66051. 37 <=> v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) [avatar definition]
66052. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) <- (~37) [avatar component clause 66051] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66053. v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) <- (37) [avatar component clause 66051] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66055. 38 <=> v1_xboole_0(sK4) [avatar definition]
66056. v1_xboole_0(sK4) <- (38) [avatar component clause 66055] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66058. 37 | ~38 [avatar split clause 45,66055,66051]
66059. 38 | ~37 [avatar split clause 44,66051,66055]
66061. 39 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) [avatar definition]
66062. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) <- (~39) [avatar component clause 66061] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66063. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) <- (39) [avatar component clause 66061] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66065. 40 <=> v1_xboole_0(sK3) [avatar definition]
66066. v1_xboole_0(sK3) <- (40) [avatar component clause 66065] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66067. ~v1_xboole_0(sK3) <- (~40) [avatar component clause 66065] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66068. 39 | ~40 [avatar split clause 43,66065,66061]
66069. 40 | ~39 [avatar split clause 42,66061,66065]
66071. 41 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) [avatar definition]
66072. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) <- (~41) [avatar component clause 66071] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66073. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) <- (41) [avatar component clause 66071] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66075. 42 <=> v1_xboole_0(sK2) [avatar definition]
66076. v1_xboole_0(sK2) <- (42) [avatar component clause 66075] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66078. 41 | ~42 [avatar split clause 41,66075,66071]
66079. 42 | ~41 [avatar split clause 40,66071,66075]
66081. 43 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) [avatar definition]
66082. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) <- (~43) [avatar component clause 66081] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66083. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) <- (43) [avatar component clause 66081] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66085. 44 <=> v1_xboole_0(sK1) [avatar definition]
66088. 43 | ~44 [avatar split clause 39,66085,66081]
66089. 44 | ~43 [avatar split clause 38,66081,66085]
66091. 45 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) [avatar definition]
66092. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) <- (~45) [avatar component clause 66091] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66093. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) <- (45) [avatar component clause 66091] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66095. 46 <=> v1_xboole_0(sK0) [avatar definition]
66096. v1_xboole_0(sK0) <- (46) [avatar component clause 66095] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66098. 45 | ~46 [avatar split clause 37,66095,66091]
66099. 46 | ~45 [avatar split clause 36,66091,66095]
66101. ~v1_xboole_0(sK17) <- (~17) [resolution 65615,65944] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66103. 47 <=> v1_xboole_0(sK17) [avatar definition]
66104. v1_xboole_0(sK17) <- (47) [avatar component clause 66103] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66105. ~v1_xboole_0(sK17) <- (~47) [avatar component clause 66103] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66106. ~47 | 17 [avatar split clause 66101,65943,66103]
66107. ~v1_xboole_0(k3_gate_1(X1,X0)) | v1_xboole_0(X1) [superposition 65614,65612] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
66110. v1_xboole_0(k2_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 65615,65613] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
66113. 48 <=> v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK18))) [avatar definition]
66114. ~v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK18))) <- (~48) [avatar component clause 66113] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66115. v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK18))) <- (48) [avatar component clause 66113] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66119. ~v1_xboole_0(sK22) <- (~48) [resolution 66114,65615] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66121. 49 <=> v1_xboole_0(sK22) [avatar definition]
66122. v1_xboole_0(sK22) <- (49) [avatar component clause 66121] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66123. ~v1_xboole_0(sK22) <- (~49) [avatar component clause 66121] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66124. ~49 | 48 [avatar split clause 66119,66113,66121]
66128. ~v1_xboole_0(k1_gate_1(sK21)) <- (~28) [resolution 65618,66004] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66129. ~v1_xboole_0(k1_gate_1(sK18)) <- (~41) [resolution 65618,66072] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66130. ~v1_xboole_0(k1_gate_1(sK18)) <- (~45) [resolution 65618,66092] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66131. ~v1_xboole_0(sK18) <- (~37) [resolution 65618,66052] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66133. 50 <=> v1_xboole_0(k1_gate_1(sK21)) [avatar definition]
66134. v1_xboole_0(k1_gate_1(sK21)) <- (50) [avatar component clause 66133] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66135. ~v1_xboole_0(k1_gate_1(sK21)) <- (~50) [avatar component clause 66133] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66136. ~50 | 28 [avatar split clause 66128,66003,66133]
66138. 51 <=> v1_xboole_0(k1_gate_1(sK18)) [avatar definition]
66139. v1_xboole_0(k1_gate_1(sK18)) <- (51) [avatar component clause 66138] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66140. ~v1_xboole_0(k1_gate_1(sK18)) <- (~51) [avatar component clause 66138] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66141. ~51 | 41 [avatar split clause 66129,66071,66138]
66142. ~51 | 45 [avatar split clause 66130,66091,66138]
66144. 52 <=> v1_xboole_0(sK18) [avatar definition]
66145. v1_xboole_0(sK18) <- (52) [avatar component clause 66144] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66146. ~v1_xboole_0(sK18) <- (~52) [avatar component clause 66144] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66147. ~52 | 37 [avatar split clause 66131,66051,66144]
66151. ~v1_xboole_0(sK20) <- (~23) [resolution 65619,65974] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66152. ~v1_xboole_0(sK20) <- (~24) [resolution 65619,65980] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66153. ~v1_xboole_0(k1_gate_1(sK20)) <- (~25) [resolution 65619,65986] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66154. ~v1_xboole_0(sK20) <- (~28) [resolution 65619,66004] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66156. ~v1_xboole_0(k1_gate_1(sK17)) <- (~45) [resolution 65619,66092] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66157. ~v1_xboole_0(k1_gate_1(sK17)) <- (~37) [resolution 65619,66052] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66167. ~20 | 28 [avatar split clause 66154,66003,65957]
66168. ~20 | 24 [avatar split clause 66152,65979,65957]
66169. ~20 | 23 [avatar split clause 66151,65973,65957]
66171. 53 <=> v1_xboole_0(k1_gate_1(sK17)) [avatar definition]
66172. v1_xboole_0(k1_gate_1(sK17)) <- (53) [avatar component clause 66171] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66173. ~v1_xboole_0(k1_gate_1(sK17)) <- (~53) [avatar component clause 66171] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66174. ~53 | 37 [avatar split clause 66157,66051,66171]
66175. ~53 | 45 [avatar split clause 66156,66091,66171]
66177. 54 <=> v1_xboole_0(k1_gate_1(sK20)) [avatar definition]
66178. v1_xboole_0(k1_gate_1(sK20)) <- (54) [avatar component clause 66177] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66179. ~v1_xboole_0(k1_gate_1(sK20)) <- (~54) [avatar component clause 66177] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66180. ~54 | 25 [avatar split clause 66153,65985,66177]
66181. ~v1_xboole_0(sK16) <- (~19) [resolution 65954,65615] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66183. 55 <=> v1_xboole_0(sK16) [avatar definition]
66184. v1_xboole_0(sK16) <- (55) [avatar component clause 66183] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66185. ~v1_xboole_0(sK16) <- (~55) [avatar component clause 66183] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66186. ~55 | 19 [avatar split clause 66181,65953,66183]
66188. ~v1_xboole_0(k1_gate_1(sK19)) <- (~24) [resolution 65620,65980] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66189. ~v1_xboole_0(sK19) <- (~25) [resolution 65620,65986] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66190. ~v1_xboole_0(k1_gate_1(sK19)) <- (~28) [resolution 65620,66004] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66191. ~v1_xboole_0(k1_gate_1(sK16)) <- (~41) [resolution 65620,66072] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66192. ~v1_xboole_0(k1_gate_1(sK16)) <- (~45) [resolution 65620,66092] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66193. ~v1_xboole_0(k1_gate_1(sK16)) <- (~37) [resolution 65620,66052] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66195. 56 <=> v1_xboole_0(k1_gate_1(sK19)) [avatar definition]
66196. v1_xboole_0(k1_gate_1(sK19)) <- (56) [avatar component clause 66195] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66197. ~v1_xboole_0(k1_gate_1(sK19)) <- (~56) [avatar component clause 66195] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66198. ~56 | 24 [avatar split clause 66188,65979,66195]
66199. ~56 | 28 [avatar split clause 66190,66003,66195]
66201. 57 <=> v1_xboole_0(k1_gate_1(sK16)) [avatar definition]
66202. v1_xboole_0(k1_gate_1(sK16)) <- (57) [avatar component clause 66201] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66203. ~v1_xboole_0(k1_gate_1(sK16)) <- (~57) [avatar component clause 66201] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66204. ~57 | 41 [avatar split clause 66191,66071,66201]
66205. ~57 | 45 [avatar split clause 66192,66091,66201]
66206. ~57 | 37 [avatar split clause 66193,66051,66201]
66215. v1_xboole_0(sK21) <- (~50) [resolution 66135,65610] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66217. $false <- (~18, ~50) [subsumption resolution 66215,65949] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66218. 18 | 50 [avatar contradiction clause 66217]
66220. v1_xboole_0(sK18) <- (~51) [resolution 66140,65610] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66222. $false <- (~51, ~52) [subsumption resolution 66220,66146] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66223. 51 | 52 [avatar contradiction clause 66222]
66225. v1_xboole_0(sK17) <- (~53) [resolution 66173,65610] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66227. $false <- (~47, ~53) [subsumption resolution 66225,66105] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66228. 47 | 53 [avatar contradiction clause 66227]
66230. v1_xboole_0(sK20) <- (~54) [resolution 66179,65610] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66234. 20 | 54 [avatar split clause 66230,66177,65957]
66236. ~v1_xboole_0(sK1) <- (~3) [resolution 65632,65615] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66240. ~44 | 3 [avatar split clause 66236,65631,66085]
66242. v1_xboole_0(sK22) | v1_xboole_0(sK16) <- (19) [resolution 65955,65616] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66247. ~22 | 25 [avatar split clause 66189,65985,65967]
66252. 55 | 49 | ~19 [avatar split clause 66242,65953,66121,66183]
66257. v1_xboole_0(sK19) <- (~56) [resolution 66197,65610] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66259. $false <- (~22, ~56) [subsumption resolution 66257,65969] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66260. 22 | 56 [avatar contradiction clause 66259]
66261. v1_xboole_0(sK16) <- (~57) [resolution 66203,65610] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66263. ~v1_xboole_0(sK16) <- (~43) [resolution 66082,65620] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66269. ~55 | 43 [avatar split clause 66263,66081,66183]
66270. 55 | 57 [avatar split clause 66261,66201,66183]
66277. ~v1_xboole_0(sK7) <- (~7) [resolution 65648,65615] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66281. ~32 | 7 [avatar split clause 66277,65647,66025]
66282. ~v1_xboole_0(sK16) <- (~31) [resolution 66022,65620] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66283. ~v1_xboole_0(sK17) <- (~31) [resolution 66022,65619] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66284. ~v1_xboole_0(sK18) <- (~31) [resolution 66022,65618] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66285. ~55 | 31 [avatar split clause 66282,66021,66183]
66286. v1_xboole_0(k1_gate_1(sK18)) | v1_xboole_0(sK22) <- (48) [resolution 66115,65616] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66287. ~v1_xboole_0(sK16) <- (57) [resolution 66202,65611] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66296. v1_xboole_0(sK16) | v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (39) [resolution 65621,66063] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66304. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (37) [resolution 65621,66053] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66305. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(sK21) <- (26) [resolution 65621,65993] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66309. v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (39, ~55) [subsumption resolution 66296,66185] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66319. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK21) <- (26, ~54) [subsumption resolution 66305,66179] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66320. v1_xboole_0(k1_gate_1(sK19)) <- (~18, 26, ~54) [subsumption resolution 66319,65949] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66321. 56 | 18 | ~26 | 54 [avatar split clause 66320,66177,65991,65947,66195]
66325. ~v1_xboole_0(sK19) <- (56) [resolution 66196,65611] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66327. $false <- (22, 56) [subsumption resolution 66325,65968] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66328. ~22 | ~56 [avatar contradiction clause 66327]
66329. ~v1_xboole_0(k1_gate_1(sK19)) <- (~26) [resolution 65992,65620] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66331. ~v1_xboole_0(sK21) <- (~26) [resolution 65992,65618] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66332. ~56 | 26 [avatar split clause 66329,65991,66195]
66333. ~v1_xboole_0(sK21) <- (50) [resolution 66134,65611] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66338. ~v1_xboole_0(sK22) <- (~7) [resolution 66110,65648] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66339. ~v1_xboole_0(sK22) <- (~17) [resolution 66110,65944] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66344. $false <- (~7, 49) [subsumption resolution 66338,66122] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66345. 7 | ~49 [avatar contradiction clause 66344]
66346. $false <- (~17, 49) [subsumption resolution 66339,66122] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66347. 17 | ~49 [avatar contradiction clause 66346]
66348. ~18 | ~50 [avatar split clause 66333,66133,65947]
66349. ~18 | 26 [avatar split clause 66331,65991,65947]
66354. ~v1_xboole_0(sK22) <- (~13) [resolution 65672,66110] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66355. ~v1_xboole_0(sK5) <- (~13) [resolution 65672,65615] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66358. $false <- (~13, 49) [subsumption resolution 66354,66122] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66359. 13 | ~49 [avatar contradiction clause 66358]
66360. $false <- (~13, 36) [subsumption resolution 66355,66046] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66361. 13 | ~36 [avatar contradiction clause 66360]
66364. 58 <=> v1_xboole_0(k3_gate_1(k1_gate_1(sK22),k2_gate_1(sK4,sK22))) [avatar definition]
66365. ~v1_xboole_0(k3_gate_1(k1_gate_1(sK22),k2_gate_1(sK4,sK22))) <- (~58) [avatar component clause 66364] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66366. v1_xboole_0(k3_gate_1(k1_gate_1(sK22),k2_gate_1(sK4,sK22))) <- (58) [avatar component clause 66364] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66368. v1_xboole_0(sK22) | v1_xboole_0(sK17) <- (17) [resolution 65945,65616] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66370. v1_xboole_0(k1_gate_1(sK22)) <- (58) [resolution 66366,66107] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66371. v1_xboole_0(k2_gate_1(sK4,sK22)) <- (58) [resolution 66366,65614] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66373. 59 <=> v1_xboole_0(k1_gate_1(sK22)) [avatar definition]
66374. ~v1_xboole_0(k1_gate_1(sK22)) <- (~59) [avatar component clause 66373] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66375. v1_xboole_0(k1_gate_1(sK22)) <- (59) [avatar component clause 66373] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66376. 59 | ~58 [avatar split clause 66370,66364,66373]
66378. 60 <=> v1_xboole_0(k2_gate_1(sK4,sK22)) [avatar definition]
66379. ~v1_xboole_0(k2_gate_1(sK4,sK22)) <- (~60) [avatar component clause 66378] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66380. v1_xboole_0(k2_gate_1(sK4,sK22)) <- (60) [avatar component clause 66378] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66381. 60 | ~58 [avatar split clause 66371,66364,66378]
66383. ~v1_xboole_0(sK22) <- (59) [resolution 66375,65611] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66385. $false <- (49, 59) [subsumption resolution 66383,66122] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66386. ~49 | ~59 [avatar contradiction clause 66385]
66387. ~v1_xboole_0(k3_gate_1(k1_gate_1(sK22),k2_gate_1(sK4,sK22))) <- (~11) [forward demodulation 65664,65612] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66388. ~58 | 11 [avatar split clause 66387,65663,66364]
66391. ~v1_xboole_0(sK22) <- (~15) [resolution 65680,66110] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66392. ~v1_xboole_0(sK2) <- (~15) [resolution 65680,65615] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66395. $false <- (~15, 49) [subsumption resolution 66391,66122] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66396. 15 | ~49 [avatar contradiction clause 66395]
66397. $false <- (~15, 42) [subsumption resolution 66392,66076] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66398. 15 | ~42 [avatar contradiction clause 66397]
66401. ~v1_xboole_0(sK22) <- (~5) [resolution 65640,66110] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66402. ~v1_xboole_0(sK3) <- (~5) [resolution 65640,65615] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66405. $false <- (~5, 49) [subsumption resolution 66401,66122] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66406. 5 | ~49 [avatar contradiction clause 66405]
66407. $false <- (~5, 40) [subsumption resolution 66402,66066] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66408. 5 | ~40 [avatar contradiction clause 66407]
66411. ~v1_xboole_0(sK22) <- (~9) [resolution 65656,66110] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66412. ~v1_xboole_0(sK6) <- (~9) [resolution 65656,65615] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66415. $false <- (~9, 49) [subsumption resolution 66411,66122] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66416. 9 | ~49 [avatar contradiction clause 66415]
66417. $false <- (~9, 34) [subsumption resolution 66412,66036] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66418. 9 | ~34 [avatar contradiction clause 66417]
66421. ~v1_xboole_0(sK19) <- (~29) [resolution 66010,65620] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66425. $false <- (22, ~29) [subsumption resolution 66421,65968] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66426. ~22 | 29 [avatar contradiction clause 66425]
66429. ~v1_xboole_0(sK19) <- (~27) [resolution 65998,65620] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66430. ~v1_xboole_0(sK20) <- (~27) [resolution 65998,65619] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66434. $false <- (22, ~27) [subsumption resolution 66429,65968] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66435. ~22 | 27 [avatar contradiction clause 66434]
66436. $false <- (20, ~27) [subsumption resolution 66430,65958] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66437. ~20 | 27 [avatar contradiction clause 66436]
66440. ~v1_xboole_0(sK22) <- (~1) [resolution 65624,66110] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66441. ~v1_xboole_0(sK0) <- (~1) [resolution 65624,65615] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66444. $false <- (~1, 49) [subsumption resolution 66440,66122] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66445. 1 | ~49 [avatar contradiction clause 66444]
66446. $false <- (~1, 46) [subsumption resolution 66441,66096] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66447. 1 | ~46 [avatar contradiction clause 66446]
66448. 51 | 47 | ~39 | 55 [avatar split clause 66309,66183,66061,66103,66138]
66449. ~47 | 31 [avatar split clause 66283,66021,66103]
66451. ~52 | 31 [avatar split clause 66284,66021,66144]
66453. ~v1_xboole_0(sK22) <- (~3) [resolution 65632,66110] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66456. $false <- (~3, 49) [subsumption resolution 66453,66122] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66457. 3 | ~49 [avatar contradiction clause 66456]
66458. v1_xboole_0(sK22) | v1_xboole_0(sK7) <- (7) [resolution 65649,65616] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66459. v1_xboole_0(sK16) | v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (31) [resolution 66023,65621] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66460. v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (31, ~55) [subsumption resolution 66459,66185] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66461. 52 | 47 | ~31 | 55 [avatar split clause 66460,66183,66021,66103,66144]
66465. ~v1_xboole_0(k1_gate_1(sK17)) <- (~43) [resolution 66082,65619] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66466. ~v1_xboole_0(k1_gate_1(sK18)) <- (~43) [resolution 66082,65618] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66467. ~53 | 43 [avatar split clause 66465,66081,66171]
66468. ~51 | 43 [avatar split clause 66466,66081,66138]
66470. ~v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(k1_gate_1(sK22)) <- (~58) [resolution 66365,65617] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66471. ~59 | ~60 | 58 [avatar split clause 66470,66364,66378,66373]
66473. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (30) [resolution 66017,65621] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66474. v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (30, ~56) [subsumption resolution 66473,66197] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66475. v1_xboole_0(k1_gate_1(sK20)) <- (30, ~50, ~56) [subsumption resolution 66474,66135] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66476. 54 | ~30 | 50 | 56 [avatar split clause 66475,66195,66133,66015,66177]
66478. ~v1_xboole_0(sK20) <- (54) [resolution 66178,65611] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66482. ~20 | ~54 [avatar split clause 66478,66177,65957]
66486. v1_xboole_0(sK22) <- (~59) [resolution 66374,65610] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66488. ~v1_xboole_0(sK22) <- (~19) [resolution 65954,66110] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66493. ~49 | 19 [avatar split clause 66488,65953,66121]
66494. 49 | 59 [avatar split clause 66486,66373,66121]
66497. 49 | 51 | ~48 [avatar split clause 66286,66113,66138,66121]
66498. v1_xboole_0(sK22) | v1_xboole_0(sK1) <- (3) [resolution 65633,65616] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66500. ~v1_xboole_0(sK18) <- (51) [resolution 66139,65611] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66504. ~52 | ~51 [avatar split clause 66500,66138,66144]
66508. v1_xboole_0(sK16) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (43) [resolution 66083,65621] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66510. ~v1_xboole_0(sK17) <- (53) [resolution 66172,65611] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66514. v1_xboole_0(sK7) <- (7, ~49) [subsumption resolution 66458,66123] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66515. 32 | ~7 | 49 [avatar split clause 66514,66121,65647,66025]
66516. ~47 | ~53 [avatar split clause 66510,66171,66103]
66517. v1_xboole_0(sK17) <- (17, ~49) [subsumption resolution 66368,66123] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66518. 47 | ~17 | 49 [avatar split clause 66517,66121,65943,66103]
66527. ~v1_xboole_0(k1_gate_1(sK17)) <- (~35) [resolution 66042,65619] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66529. ~53 | 35 [avatar split clause 66527,66041,66171]
66536. ~v1_xboole_0(sK4) <- (~60) [resolution 66379,65615] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66538. $false <- (38, ~60) [subsumption resolution 66536,66056] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66539. ~38 | 60 [avatar contradiction clause 66538]
66540. v1_xboole_0(k3_gate_1(k1_gate_1(sK22),k2_gate_1(sK4,sK22))) <- (11) [forward demodulation 65665,65612] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66541. 58 | ~11 [avatar split clause 66540,65663,66364]
66544. v1_xboole_0(sK22) | v1_xboole_0(sK5) <- (13) [resolution 65673,65616] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66546. v1_xboole_0(sK5) <- (13, ~49) [subsumption resolution 66544,66123] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66547. $false <- (13, ~36, ~49) [subsumption resolution 66546,66047] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66548. ~13 | 36 | 49 [avatar contradiction clause 66547]
66553. ~v1_xboole_0(k1_gate_1(sK16)) <- (~33) [resolution 66032,65620] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66554. ~v1_xboole_0(sK17) <- (~33) [resolution 66032,65619] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66558. $false <- (~33, 57) [subsumption resolution 66553,66202] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66559. 33 | ~57 [avatar contradiction clause 66558]
66560. $false <- (~33, 47) [subsumption resolution 66554,66104] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66561. 33 | ~47 [avatar contradiction clause 66560]
66565. ~v1_xboole_0(sK19) <- (~23) [resolution 65974,65620] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66567. ~v1_xboole_0(sK21) <- (~23) [resolution 65974,65618] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66570. $false <- (22, ~23) [subsumption resolution 66565,65968] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66571. ~22 | 23 [avatar contradiction clause 66570]
66572. $false <- (18, ~23) [subsumption resolution 66567,65948] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66573. ~18 | 23 [avatar contradiction clause 66572]
66576. v1_xboole_0(sK22) | v1_xboole_0(sK4) <- (60) [resolution 66380,65616] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66578. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (33) [resolution 66033,65621] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66582. ~v1_xboole_0(sK17) <- (~39) [resolution 66062,65619] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66583. ~v1_xboole_0(k1_gate_1(sK18)) <- (~39) [resolution 66062,65618] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66586. $false <- (~39, 47) [subsumption resolution 66582,66104] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66587. 39 | ~47 [avatar contradiction clause 66586]
66588. $false <- (~39, 51) [subsumption resolution 66583,66139] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66589. 39 | ~51 [avatar contradiction clause 66588]
66593. ~v1_xboole_0(sK21) <- (~24) [resolution 65980,65618] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66595. $false <- (18, ~24) [subsumption resolution 66593,65948] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66596. ~18 | 24 [avatar contradiction clause 66595]
66603. ~v1_xboole_0(k1_gate_1(sK20)) <- (~30) [resolution 66016,65619] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66606. $false <- (~30, 54) [subsumption resolution 66603,66178] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66607. 30 | ~54 [avatar contradiction clause 66606]
66610. v1_xboole_0(sK22) | v1_xboole_0(sK6) <- (9) [resolution 65657,65616] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66621. v1_xboole_0(sK16) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (35) [resolution 66043,65621] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66624. ~v1_xboole_0(k1_gate_1(sK20)) <- (~26) [resolution 65992,65619] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66627. $false <- (~26, 54) [subsumption resolution 66624,66178] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66628. 26 | ~54 [avatar contradiction clause 66627]
66638. ~55 | ~57 [avatar split clause 66287,66201,66183]
66652. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (24) [resolution 65981,65621] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66657. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (28) [resolution 66005,65621] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66661. ~v1_xboole_0(k1_gate_1(sK20)) <- (~29) [resolution 66010,65619] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66662. ~v1_xboole_0(k1_gate_1(sK21)) <- (~29) [resolution 66010,65618] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66665. $false <- (~29, 54) [subsumption resolution 66661,66178] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66666. 29 | ~54 [avatar contradiction clause 66665]
66667. $false <- (~29, 50) [subsumption resolution 66662,66134] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66668. 29 | ~50 [avatar contradiction clause 66667]
66670. v1_xboole_0(sK22) | v1_xboole_0(sK3) <- (5) [resolution 65641,65616] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66672. v1_xboole_0(sK3) <- (5, ~49) [subsumption resolution 66670,66123] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66673. $false <- (5, ~40, ~49) [subsumption resolution 66672,66067] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66674. ~5 | 40 | 49 [avatar contradiction clause 66673]
66678. ~v1_xboole_0(k1_gate_1(sK21)) <- (~27) [resolution 65998,65618] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66680. $false <- (~27, 50) [subsumption resolution 66678,66134] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66681. 27 | ~50 [avatar contradiction clause 66680]
66685. 51 | 53 | 55 | ~43 [avatar split clause 66508,66081,66183,66171,66138]
66686. 52 | 53 | 55 | ~35 [avatar split clause 66621,66041,66183,66171,66144]
66687. 52 | 47 | 57 | ~33 [avatar split clause 66578,66031,66201,66103,66144]
66688. 52 | 53 | 57 | ~37 [avatar split clause 66304,66051,66201,66171,66144]
66695. v1_xboole_0(sK19) | v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (23) [resolution 65975,65621] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66702. v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (29) [resolution 66011,65621] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66703. v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (~22, 29) [subsumption resolution 66702,65969] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66704. v1_xboole_0(k1_gate_1(sK20)) <- (~22, 29, ~50) [subsumption resolution 66703,66135] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66705. 54 | 22 | ~29 | 50 [avatar split clause 66704,66133,66009,65967,66177]
66707. v1_xboole_0(sK22) | v1_xboole_0(sK0) <- (1) [resolution 65625,65616] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66709. v1_xboole_0(sK0) <- (1, ~49) [subsumption resolution 66707,66123] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66712. 46 | ~1 | 49 [avatar split clause 66709,66121,65623,66095]
66720. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (45) [resolution 66093,65621] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66722. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (45, ~57) [subsumption resolution 66720,66203] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66723. v1_xboole_0(k1_gate_1(sK18)) <- (45, ~53, ~57) [subsumption resolution 66722,66173] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66724. $false <- (45, ~51, ~53, ~57) [subsumption resolution 66723,66140] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66725. ~45 | 51 | 53 | 57 [avatar contradiction clause 66724]
66732. ~v1_xboole_0(sK17) <- (~41) [resolution 66072,65619] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66735. $false <- (~41, 47) [subsumption resolution 66732,66104] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66736. 41 | ~47 [avatar contradiction clause 66735]
66741. ~v1_xboole_0(sK16) <- (~35) [resolution 66042,65620] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66743. ~v1_xboole_0(sK18) <- (~35) [resolution 66042,65618] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66746. $false <- (~35, 55) [subsumption resolution 66741,66184] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66747. 35 | ~55 [avatar contradiction clause 66746]
66748. $false <- (~35, 52) [subsumption resolution 66743,66145] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66749. 35 | ~52 [avatar contradiction clause 66748]
66750. v1_xboole_0(sK22) | v1_xboole_0(sK2) <- (15) [resolution 65681,65616] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66754. ~v1_xboole_0(sK21) <- (~25) [resolution 65986,65618] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66756. $false <- (18, ~25) [subsumption resolution 66754,65948] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66757. ~18 | 25 [avatar contradiction clause 66756]
66758. v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(sK21) <- (25) [resolution 65987,65621] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66760. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (41) [resolution 66073,65621] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66762. ~v1_xboole_0(k1_gate_1(sK19)) <- (~30) [resolution 66016,65620] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66764. ~v1_xboole_0(k1_gate_1(sK21)) <- (~30) [resolution 66016,65618] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66768. v1_xboole_0(sK4) <- (~49, 60) [subsumption resolution 66576,66123] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66769. 38 | 49 | ~60 [avatar split clause 66768,66378,66121,66055]
66770. ~56 | 30 [avatar split clause 66762,66015,66195]
66771. v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK18))) <- (21) [forward demodulation 65965,65613] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66772. 48 | ~21 [avatar split clause 66771,65963,66113]
66773. v1_xboole_0(sK19) | v1_xboole_0(sK21) <- (25, ~54) [subsumption resolution 66758,66179] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66774. 18 | 22 | ~25 | 54 [avatar split clause 66773,66177,65985,65967,65947]
66775. v1_xboole_0(sK6) <- (9, ~49) [subsumption resolution 66610,66123] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66776. 34 | ~9 | 49 [avatar split clause 66775,66121,65655,66035]
66777. v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (41, ~57) [subsumption resolution 66760,66203] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66778. 51 | 47 | ~41 | 57 [avatar split clause 66777,66201,66071,66103,66138]
66779. ~50 | 30 [avatar split clause 66764,66015,66133]
66780. v1_xboole_0(sK2) <- (15, ~49) [subsumption resolution 66750,66123] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66781. 42 | ~15 | 49 [avatar split clause 66780,66121,65679,66075]
66782. ~v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK18))) <- (~21) [forward demodulation 65964,65613] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66783. ~48 | 21 [avatar split clause 66782,65963,66113]
66790. ~v1_xboole_0(k1_gate_1(sK18)) <- (~48) [resolution 66114,66110] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66792. ~51 | 48 [avatar split clause 66790,66113,66138]
66801. ~v1_xboole_0(sK18) <- (~33) [resolution 66032,65618] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66803. $false <- (~33, 52) [subsumption resolution 66801,66145] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66804. 33 | ~52 [avatar contradiction clause 66803]
66810. ~v1_xboole_0(sK16) <- (~39) [resolution 66062,65620] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66814. $false <- (~39, 55) [subsumption resolution 66810,66184] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66815. 39 | ~55 [avatar contradiction clause 66814]
66816. 18 | 20 | 56 | ~24 [avatar split clause 66652,65979,66195,65957,65947]
66817. 50 | 20 | 56 | ~28 [avatar split clause 66657,66003,66195,65957,66133]
66818. v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (~22, 23) [subsumption resolution 66695,65969] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66819. 18 | 20 | 22 | ~23 [avatar split clause 66818,65973,65967,65957,65947]
66820. v1_xboole_0(sK1) <- (3, ~49) [subsumption resolution 66498,66123] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66821. 44 | ~3 | 49 [avatar split clause 66820,66121,65631,66085]
66830. v1_xboole_0(sK19) | v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (27) [resolution 65999,65621] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66832. v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (~22, 27) [subsumption resolution 66830,65969] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66833. v1_xboole_0(k1_gate_1(sK21)) <- (~20, ~22, 27) [subsumption resolution 66832,65959] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66834. $false <- (~20, ~22, 27, ~50) [subsumption resolution 66833,66135] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66835. 20 | 22 | ~27 | 50 [avatar contradiction clause 66834]
66836. $false [avatar sat refutation 65813,65877,65909,65925,65933,65937,65939,65940,65950,65951,65960,65961,65970,65971,65976,65977,65982,65983,65988,65989,65994,65995,66000,66001,66006,66007,66012,66013,66018,66019,66028,66029,66038,66039,66048,66049,66058,66059,66068,66069,66078,66079,66088,66089,66098,66099,66106,66124,66136,66141,66142,66147,66167,66168,66169,66174,66175,66180,66186,66198,66199,66204,66205,66206,66218,66223,66228,66234,66240,66247,66252,66260,66269,66270,66281,66285,66321,66328,66332,66345,66347,66348,66349,66359,66361,66376,66381,66386,66388,66396,66398,66406,66408,66416,66418,66426,66435,66437,66445,66447,66448,66449,66451,66457,66461,66467,66468,66471,66476,66482,66493,66494,66497,66504,66515,66516,66518,66529,66539,66541,66548,66559,66561,66571,66573,66587,66589,66596,66607,66628,66638,66666,66668,66674,66681,66685,66686,66687,66688,66705,66712,66725,66736,66747,66749,66757,66769,66770,66772,66774,66776,66778,66779,66781,66783,66792,66804,66815,66816,66817,66819,66821,66835] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:5876,thDist:-5876}
% SZS output end Proof for gate_3__t4_gate_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 23
% 
% >>> Saturation
% Initial clauses: 65586
% Generated clauses: 66524
% Active clauses: 330
% Passive clauses: 407
% Final active clauses: 55
% Final passive clauses: 13
% Main loop iterations started: 331
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 100
% Fw demodulations: 10
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 65288
% Forward subsumptions: 207
% 
% >>> Generating Inferences
% Binary resolution: 284
% Unit resulting resolution: 59
% Forward superposition: 8
% Backward superposition: 4
% 
% >>> AVATAR
% Split clauses: 315
% Split components: 4221
% Sat splitting refutations: 118
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 510
% SAT solver binary clauses: 202
% 
% Memory used [KB]: 31598
% Time elapsed: 0.146 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.068 s
% preprocessing: 0.012 s
% SAT solver time: 0.002 s
% minimizing solver time: 0.001 s
% congruence closure: 0.003 s ( own 0.002 s ) 
% backward demodulation index maintenance: 0.001 s
% forward subsumption: 0.001 s
% forward demodulation: 0.001 s
% unit resulting resolution: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.146 s ( own 0.056 s ) 

