i: [1,22,0,3,1,1,0,0] t35_xboole_1
a: [3,36,0,3,1,22]
s: 36
i: [1,21,0,3,1,1,0,0] t35_xboole_1
a: [3,41,0,3,1,21]
s: 41
i: [1,23,0,7,1,1,0,0] t35_xboole_1
a: [3,31,0,7,1,23]
s: 31
i: [1,27,0,3,1,0,0,1] reflexivity_r1_tarski
s: 27
i: [1,26,0,9,3,0,0,1] t1_xboole_1
s: 26
i: [1,24,0,10,2,0,0,1] t34_xboole_1
s: 24
d: [2,43,1,6,2,1,40,26,36]
s: 43
i: [1,25,0,10,2,0,0,1] t33_xboole_1
s: 25
d: [2,44,1,6,2,1,40,26,41]
s: 44
d: [2,46,1,13,3,0,40,24,26]
s: 46
d: [2,48,1,13,3,0,40,25,26]
s: 48
d: [2,51,2,13,3,0,40,46,24]
s: 51
d: [2,52,2,13,3,0,40,46,25]
s: 52
d: [2,59,3,0,0,3,52,41,36,31,52]
e: 59
d: [2,60,3,6,2,1,40,52,31]
d: [2,65,3,3,1,2,21,60,36]
d: [2,66,3,0,0,3,21,65,41]
e: 66
f: 66,21,22,23
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t35_xboole_1
% SZS output start Proof for xboole_1__t35_xboole_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => r1_tarski(k4_xboole_0(X0,X3),k4_xboole_0(X1,X2))) [input t35_xboole_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => r1_tarski(k4_xboole_0(X0,X3),k4_xboole_0(X1,X2))) [negated conjecture 1]
5. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
6. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,X1) => r1_tarski(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2))) [input t33_xboole_1]
7. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,X1) => r1_tarski(k4_xboole_0(X2,X1),k4_xboole_0(X2,X0))) [input t34_xboole_1]
8. ~! [X0,X1,X2,X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => r1_tarski(k4_xboole_0(X0,X3),k4_xboole_0(X1,X2))) [flattening 2]
9. ! [X0,X1,X2] : (r1_tarski(X0,X1) => r1_tarski(k4_xboole_0(X2,X1),k4_xboole_0(X2,X0))) [flattening 7]
10. ! [X0,X1,X2] : (r1_tarski(X0,X1) => r1_tarski(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2))) [flattening 6]
11. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 5]
13. ? [X0,X1,X2,X3] : (~r1_tarski(k4_xboole_0(X0,X3),k4_xboole_0(X1,X2)) & (r1_tarski(X2,X3) & r1_tarski(X0,X1))) [ennf transformation 8]
14. ? [X0,X1,X2,X3] : (~r1_tarski(k4_xboole_0(X0,X3),k4_xboole_0(X1,X2)) & r1_tarski(X2,X3) & r1_tarski(X0,X1)) [flattening 13]
15. ! [X0,X1,X2] : (r1_tarski(k4_xboole_0(X2,X1),k4_xboole_0(X2,X0)) | ~r1_tarski(X0,X1)) [ennf transformation 9]
16. ! [X0,X1,X2] : (r1_tarski(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2)) | ~r1_tarski(X0,X1)) [ennf transformation 10]
17. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 11]
18. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 17]
19. ? [X0,X1,X2,X3] : (~r1_tarski(k4_xboole_0(X0,X3),k4_xboole_0(X1,X2)) & r1_tarski(X2,X3) & r1_tarski(X0,X1)) => (~r1_tarski(k4_xboole_0(sK0,sK3),k4_xboole_0(sK1,sK2)) & r1_tarski(sK2,sK3) & r1_tarski(sK0,sK1)) [choice axiom]
20. ~r1_tarski(k4_xboole_0(sK0,sK3),k4_xboole_0(sK1,sK2)) & r1_tarski(sK2,sK3) & r1_tarski(sK0,sK1) [skolemisation 14,19]
21. r1_tarski(sK0,sK1) [cnf transformation 20] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
22. r1_tarski(sK2,sK3) [cnf transformation 20] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
23. ~r1_tarski(k4_xboole_0(sK0,sK3),k4_xboole_0(sK1,sK2)) [cnf transformation 20] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
24. r1_tarski(k4_xboole_0(X2,X1),k4_xboole_0(X2,X0)) | ~r1_tarski(X0,X1) [cnf transformation 15] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
25. r1_tarski(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2)) | ~r1_tarski(X0,X1) [cnf transformation 16] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
26. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 18] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
29. 1 <=> r1_tarski(k4_xboole_0(sK0,sK3),k4_xboole_0(sK1,sK2)) [avatar definition]
31. ~r1_tarski(k4_xboole_0(sK0,sK3),k4_xboole_0(sK1,sK2)) <- (~1) [avatar component clause 29] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
32. ~1 [avatar split clause 23,29]
34. 2 <=> r1_tarski(sK2,sK3) [avatar definition]
36. r1_tarski(sK2,sK3) <- (2) [avatar component clause 34] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
37. 2 [avatar split clause 22,34]
39. 3 <=> r1_tarski(sK0,sK1) [avatar definition]
41. r1_tarski(sK0,sK1) <- (3) [avatar component clause 39] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
42. 3 [avatar split clause 21,39]
46. ~r1_tarski(X2,k4_xboole_0(X3,X1)) | r1_tarski(X2,k4_xboole_0(X3,X0)) | ~r1_tarski(X0,X1) [resolution 24,26] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
52. r1_tarski(k4_xboole_0(X7,X8),k4_xboole_0(X9,X10)) | ~r1_tarski(X10,X8) | ~r1_tarski(X7,X9) [resolution 46,25] {a:2,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
60. ~r1_tarski(sK2,sK3) | ~r1_tarski(sK0,sK1) <- (~1) [resolution 52,31] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
65. ~r1_tarski(sK0,sK1) <- (~1, 2) [subsumption resolution 60,36] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66. $false <- (~1, 2, 3) [subsumption resolution 65,41] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
67. 1 | ~2 | ~3 [avatar contradiction clause 66]
68. $false [avatar sat refutation 32,37,42,67] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:23,thDist:-23}
% SZS output end Proof for xboole_1__t35_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 6
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 7
% Generated clauses: 33
% Active clauses: 13
% Passive clauses: 18
% Final active clauses: 13
% Final passive clauses: 5
% Main loop iterations started: 13
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 7
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

