i: [1,39,0,2,1,1,0,0] t40_int_1
a: [3,62,0,2,1,39]
s: 62
i: [1,40,0,6,1,1,0,0] t40_int_1
a: [3,57,0,6,1,40]
s: 57
i: [1,52,0,2,1,0,0,1] rc1_xreal_0
a: [3,67,0,2,1,52]
s: 67
i: [1,53,0,2,1,0,0,2] rc2_int_1
a: [3,72,0,2,1,53]
s: 72
i: [1,51,0,4,2,0,0,2] cc3_int_1
s: 51
i: [1,41,0,5,2,0,0,1] dt_k2_int_1
s: 41
d: [2,74,1,2,1,1,40,51,72]
a: [3,78,1,2,1,74]
s: 78
i: [1,42,0,5,2,0,0,1] dt_k2_int_1
s: 42
i: [1,43,0,5,2,0,0,1] dt_k1_int_1
s: 43
i: [1,44,0,5,2,0,0,1] dt_k1_int_1
s: 44
i: [1,45,0,8,2,0,0,1] projectivity_k2_int_1
s: 45
d: [2,80,1,5,2,0,40,41,51]
s: 80
i: [1,46,0,8,2,0,0,1] projectivity_k2_int_1
s: 46
d: [2,81,1,5,2,0,40,43,51]
s: 81
i: [1,47,0,8,2,0,0,1] projectivity_k1_int_1
s: 47
i: [1,48,0,8,2,0,0,1] projectivity_k1_int_1
s: 48
i: [1,49,0,8,3,0,0,1] t30_int_1
s: 49
i: [1,50,0,8,3,0,0,1] t30_int_1
s: 50
d: [2,90,0,6,2,0,21,50,51]
s: 90
d: [2,91,1,8,2,1,44,57,90]
d: [2,96,1,3,1,1,19,91]
a: [3,100,1,3,1,96]
s: 100
d: [2,102,2,0,0,2,52,62,100,43]
e: 102
d: [2,103,2,2,1,1,40,100,43]
d: [2,105,2,0,0,2,21,103,62]
e: 105
f: 105,96,39,40
% Refutation found. Thanks to Tanya!
% SZS status Theorem for int_1__t40_int_1
% SZS output start Proof for int_1__t40_int_1
1. ! [X0] : (v1_xreal_0(X0) => k1_int_1(X0) = k2_int_1(k1_int_1(X0))) [input t40_int_1]
2. ~! [X0] : (v1_xreal_0(X0) => k1_int_1(X0) = k2_int_1(k1_int_1(X0))) [negated conjecture 1]
8. ! [X0] : (v1_xreal_0(X0) => v1_int_1(k1_int_1(X0))) [input dt_k1_int_1]
12. ! [X0] : (v1_int_1(X0) => v1_xreal_0(X0)) [input cc3_int_1]
17. ! [X0] : (v1_xreal_0(X0) => (k2_int_1(X0) = X0 <=> v1_int_1(X0))) [input t30_int_1]
21. ? [X0] : (k1_int_1(X0) != k2_int_1(k1_int_1(X0)) & v1_xreal_0(X0)) [ennf transformation 2]
24. ! [X0] : (v1_int_1(k1_int_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 8]
30. ! [X0] : ((k2_int_1(X0) = X0 <=> v1_int_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 17]
31. ! [X0] : (v1_xreal_0(X0) | ~v1_int_1(X0)) [ennf transformation 12]
32. ? [X0] : (k1_int_1(X0) != k2_int_1(k1_int_1(X0)) & v1_xreal_0(X0)) => (k1_int_1(sK0) != k2_int_1(k1_int_1(sK0)) & v1_xreal_0(sK0)) [choice axiom]
33. k1_int_1(sK0) != k2_int_1(k1_int_1(sK0)) & v1_xreal_0(sK0) [skolemisation 21,32]
34. ! [X0] : (((k2_int_1(X0) = X0 | ~v1_int_1(X0)) & (v1_int_1(X0) | k2_int_1(X0) != X0)) | ~v1_xreal_0(X0)) [nnf transformation 30]
39. v1_xreal_0(sK0) [cnf transformation 33] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. k1_int_1(sK0) != k2_int_1(k1_int_1(sK0)) [cnf transformation 33] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. v1_int_1(k1_int_1(X0)) | ~v1_xreal_0(X0) [cnf transformation 24] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. k2_int_1(X0) = X0 | ~v1_int_1(X0) | ~v1_xreal_0(X0) [cnf transformation 34] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
51. ~v1_int_1(X0) | v1_xreal_0(X0) [cnf transformation 31] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
55. 1 <=> k1_int_1(sK0) = k2_int_1(k1_int_1(sK0)) [avatar definition]
57. k1_int_1(sK0) != k2_int_1(k1_int_1(sK0)) <- (~1) [avatar component clause 55] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
58. ~1 [avatar split clause 40,55]
60. 2 <=> v1_xreal_0(sK0) [avatar definition]
62. v1_xreal_0(sK0) <- (2) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. 2 [avatar split clause 39,60]
90. k2_int_1(X0) = X0 | ~v1_int_1(X0) [subsumption resolution 50,51] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
91. k1_int_1(sK0) != k1_int_1(sK0) | ~v1_int_1(k1_int_1(sK0)) <- (~1) [superposition 57,90] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
96. ~v1_int_1(k1_int_1(sK0)) <- (~1) [trivial inequality removal 91] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
98. 6 <=> v1_int_1(k1_int_1(sK0)) [avatar definition]
100. ~v1_int_1(k1_int_1(sK0)) <- (~6) [avatar component clause 98] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. ~6 | 1 [avatar split clause 96,55,98]
103. ~v1_xreal_0(sK0) <- (~6) [resolution 100,43] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
105. $false <- (2, ~6) [subsumption resolution 103,62] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
106. ~2 | 6 [avatar contradiction clause 105]
107. $false [avatar sat refutation 58,63,101,106] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:32,thDist:-32}
% SZS output end Proof for int_1__t40_int_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% >>> Preprocessing
% Introduced skolems: 3
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 42
% Active clauses: 15
% Passive clauses: 20
% Final active clauses: 15
% Main loop iterations started: 20
% 
% >>> Simplifying Inferences
% Trivial inequalities: 2
% Fw subsumption resolutions: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 8
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 4
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

