i: [1,60,0,2,1,1,0,0] t2_topmetr2
a: [3,99,0,2,1,60]
s: 99
i: [1,59,0,2,1,1,0,0] t2_topmetr2
a: [3,104,0,2,1,59]
s: 104
i: [1,58,0,2,1,1,0,0] t2_topmetr2
a: [3,109,0,2,1,58]
s: 109
i: [1,57,0,2,1,1,0,0] t2_topmetr2
a: [3,114,0,2,1,57]
s: 114
i: [1,62,0,10,1,1,0,0] t2_topmetr2
a: [3,89,0,10,1,62]
s: 89
i: [1,61,0,10,1,1,0,0] t2_topmetr2
a: [3,94,0,10,1,61]
s: 94
i: [1,74,0,3,1,0,0,1] reflexivity_r1_tarski
s: 74
i: [1,84,0,3,1,0,0,1] d10_xboole_0
s: 84
i: [1,85,0,3,1,0,0,1] d10_xboole_0
s: 85
i: [1,73,0,4,1,0,0,255] existence_m1_subset_1
s: 73
i: [1,75,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 75
i: [1,76,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 76
i: [1,77,0,6,1,0,0,255] dt_k6_subset_1
s: 77
i: [1,70,0,7,2,0,0,255] t3_subset
s: 70
i: [1,71,0,7,2,0,0,255] t3_subset
s: 71
i: [1,78,0,7,1,0,0,255] redefinition_k6_subset_1
s: 78
d: [2,123,1,5,1,0,40,70,73]
s: 123
i: [1,79,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 79
d: [2,124,1,5,1,0,40,70,77]
s: 124
i: [1,80,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 80
i: [1,66,0,9,1,0,0,255] t39_xboole_1
s: 66
i: [1,69,0,9,3,0,0,1] d10_xboole_0
s: 69
i: [1,81,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,121,0,9,3,81]
s: 121
d: [2,128,1,20,2,1,40,69,94]
a: [3,134,1,10,1,128]
s: 134
d: [2,126,0,9,1,0,22,66,78]
s: 126
i: [1,72,0,11,3,0,0,1] t8_xboole_1
s: 72
d: [2,129,2,10,2,0,40,69,124]
s: 129
i: [1,82,0,12,5,0,0,1] dt_k1_funct_4
s: 82
d: [2,130,2,10,2,0,40,69,123]
s: 130
i: [1,83,0,12,5,0,0,1] dt_k1_funct_4
s: 83
i: [1,63,0,15,5,0,0,1] t18_funct_4
s: 63
d: [2,144,1,13,3,0,44,72,126]
s: 144
i: [1,64,0,18,5,0,0,1] t17_funct_4
s: 64
d: [2,155,2,8,2,0,40,144,124]
s: 155
i: [1,65,0,28,6,0,0,1] l2_topmetr2
s: 65
d: [2,170,1,18,5,1,40,65,94]
d: [2,175,1,16,4,2,21,170,99]
d: [2,176,1,14,3,3,21,175,104]
d: [2,177,1,12,2,4,21,176,109]
d: [2,178,1,10,1,5,21,177,114]
a: [3,182,1,10,1,178]
s: 182
d: [2,184,2,17,2,1,40,182,144]
d: [2,186,2,17,2,1,22,184,80]
a: [3,190,2,7,1,186]
s: 190
d: [2,205,3,0,0,5,52,104,99,114,109,190,63]
e: 205
d: [2,206,3,8,4,1,40,190,63]
d: [2,208,3,6,3,2,21,206,109]
d: [2,209,3,4,2,3,21,208,114]
d: [2,210,3,2,1,4,21,209,99]
d: [2,211,3,0,0,5,21,210,104]
e: 211
r: 190
d: [2,159,1,15,5,0,44,64,76]
d: [2,162,1,11,3,0,18,159]
s: 162
a: [3,189,2,7,1,186]
s: 189
d: [2,140,1,16,4,0,40,72,69]
s: 140
d: [2,160,1,18,5,0,44,64,80]
s: 160
d: [2,217,3,14,2,1,40,189,69]
a: [3,221,3,7,1,217]
s: 221
d: [2,161,1,18,5,0,44,64,80]
s: 161
a: [3,194,2,10,1,186]
s: 194
d: [2,171,1,25,6,0,44,65,75]
d: [2,174,1,21,4,0,18,171]
s: 174
d: [2,152,1,22,6,0,40,63,69]
s: 152
d: [2,185,2,20,2,1,40,182,69]
a: [3,199,2,10,1,185]
s: 199
d: [2,157,1,28,6,0,40,64,69]
s: 157
d: [2,265,3,0,0,6,52,99,104,114,109,89,194,157]
e: 265
d: [2,266,3,18,5,1,40,157,194]
d: [2,275,3,16,4,2,21,266,104]
d: [2,276,3,14,3,3,21,275,109]
d: [2,277,3,12,2,4,21,276,114]
d: [2,278,3,2,1,5,21,277,89]
d: [2,279,3,0,0,6,21,278,99]
e: 279
f: 279,211,186,178,57,58,59,60,61,62
% Refutation found. Thanks to Tanya!
% SZS status Theorem for topmetr2__t2_topmetr2
% SZS output start Proof for topmetr2__t2_topmetr2
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(k7_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))),k10_xtuple_0(X1)) => k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) = k10_xtuple_0(k1_funct_4(X0,X1))))) [input t2_topmetr2]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(k7_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))),k10_xtuple_0(X1)) => k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) = k10_xtuple_0(k1_funct_4(X0,X1))))) [negated conjecture 1]
3. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
5. ! [X0] : ! [X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [input d10_xboole_0]
20. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(k7_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))),k10_xtuple_0(X1)) => r1_tarski(k6_subset_1(k10_xtuple_0(X0),k10_xtuple_0(X1)),k10_xtuple_0(k1_funct_4(X0,X1)))))) [input l2_topmetr2]
21. ! [X0,X1] : k6_subset_1(X0,X1) = k4_xboole_0(X0,X1) [input redefinition_k6_subset_1]
23. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => r1_tarski(k10_xtuple_0(k1_funct_4(X0,X1)),k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1))))) [input t17_funct_4]
24. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(k1_funct_4(X1,X0))))) [input t18_funct_4]
25. ! [X0] : ! [X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X0,k4_xboole_0(X1,X0)) [input t39_xboole_1]
27. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [input t8_xboole_1]
28. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X0,k4_xboole_0(X1,X0)) [flattening 25]
29. ! [X0,X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [flattening 5]
31. ! [X0,X1,X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [flattening 27]
35. ? [X0] : (? [X1] : ((k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) != k10_xtuple_0(k1_funct_4(X0,X1)) & r1_tarski(k7_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))),k10_xtuple_0(X1))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
36. ? [X0] : (? [X1] : (k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) != k10_xtuple_0(k1_funct_4(X0,X1)) & r1_tarski(k7_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))),k10_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 35]
37. ! [X0] : (! [X1] : (r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(k1_funct_4(X1,X0))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 24]
38. ! [X0] : (! [X1] : (r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(k1_funct_4(X1,X0))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 37]
39. ! [X0] : (! [X1] : (r1_tarski(k10_xtuple_0(k1_funct_4(X0,X1)),k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 23]
40. ! [X0] : (! [X1] : (r1_tarski(k10_xtuple_0(k1_funct_4(X0,X1)),k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 39]
41. ! [X0] : (! [X1] : ((r1_tarski(k6_subset_1(k10_xtuple_0(X0),k10_xtuple_0(X1)),k10_xtuple_0(k1_funct_4(X0,X1))) | ~r1_tarski(k7_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))),k10_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 20]
42. ! [X0] : (! [X1] : (r1_tarski(k6_subset_1(k10_xtuple_0(X0),k10_xtuple_0(X1)),k10_xtuple_0(k1_funct_4(X0,X1))) | ~r1_tarski(k7_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))),k10_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 41]
43. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | (~r1_tarski(X2,X1) | ~r1_tarski(X0,X1))) [ennf transformation 31]
44. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1)) [flattening 43]
49. ? [X0] : (? [X1] : (k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) != k10_xtuple_0(k1_funct_4(X0,X1)) & r1_tarski(k7_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))),k10_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(X1)) != k10_xtuple_0(k1_funct_4(sK0,X1)) & r1_tarski(k7_relat_1(sK0,k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1))),k10_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
50. ? [X1] : (k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(X1)) != k10_xtuple_0(k1_funct_4(sK0,X1)) & r1_tarski(k7_relat_1(sK0,k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1))),k10_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) => (k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) != k10_xtuple_0(k1_funct_4(sK0,sK1)) & r1_tarski(k7_relat_1(sK0,k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))),k10_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
51. (k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) != k10_xtuple_0(k1_funct_4(sK0,sK1)) & r1_tarski(k7_relat_1(sK0,k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))),k10_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 36,50,49]
52. ! [X0,X1] : ((X0 = X1 | (~r1_tarski(X1,X0) | ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [nnf transformation 29]
53. ! [X0,X1] : ((X0 = X1 | ~r1_tarski(X1,X0) | ~r1_tarski(X0,X1)) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [flattening 52]
57. v1_relat_1(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. v1_funct_1(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
59. v1_relat_1(sK1) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
60. v1_funct_1(sK1) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
61. r1_tarski(k7_relat_1(sK0,k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))),k10_xtuple_0(sK1)) [cnf transformation 51] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
62. k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) != k10_xtuple_0(k1_funct_4(sK0,sK1)) [cnf transformation 51] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
63. r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(k1_funct_4(X1,X0))) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 38] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. r1_tarski(k10_xtuple_0(k1_funct_4(X0,X1)),k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 40] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. ~r1_tarski(k7_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))),k10_xtuple_0(X1)) | r1_tarski(k6_subset_1(k10_xtuple_0(X0),k10_xtuple_0(X1)),k10_xtuple_0(k1_funct_4(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 42] {a:0,w:28,wCS:11200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
66. k2_xboole_0(X0,X1) = k2_xboole_0(X0,k4_xboole_0(X1,X0)) [cnf transformation 28] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
69. ~r1_tarski(X1,X0) | X0 = X1 | ~r1_tarski(X0,X1) [cnf transformation 53] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1) [cnf transformation 44] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
78. k6_subset_1(X0,X1) = k4_xboole_0(X0,X1) [cnf transformation 21] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
80. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. 1 <=> k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) = k10_xtuple_0(k1_funct_4(sK0,sK1)) [avatar definition]
89. k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) != k10_xtuple_0(k1_funct_4(sK0,sK1)) <- (~1) [avatar component clause 87] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
90. ~1 [avatar split clause 62,87]
92. 2 <=> r1_tarski(k7_relat_1(sK0,k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))),k10_xtuple_0(sK1)) [avatar definition]
94. r1_tarski(k7_relat_1(sK0,k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))),k10_xtuple_0(sK1)) <- (2) [avatar component clause 92] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. 2 [avatar split clause 61,92]
97. 3 <=> v1_funct_1(sK1) [avatar definition]
99. v1_funct_1(sK1) <- (3) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. 3 [avatar split clause 60,97]
102. 4 <=> v1_relat_1(sK1) [avatar definition]
104. v1_relat_1(sK1) <- (4) [avatar component clause 102] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 4 [avatar split clause 59,102]
107. 5 <=> v1_funct_1(sK0) [avatar definition]
109. v1_funct_1(sK0) <- (5) [avatar component clause 107] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. 5 [avatar split clause 58,107]
112. 6 <=> v1_relat_1(sK0) [avatar definition]
114. v1_relat_1(sK0) <- (6) [avatar component clause 112] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
115. 6 [avatar split clause 57,112]
126. k2_xboole_0(X0,X1) = k2_xboole_0(X0,k6_subset_1(X1,X0)) [forward demodulation 66,78] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
144. ~r1_tarski(k6_subset_1(X9,X8),X10) | r1_tarski(k2_xboole_0(X8,X9),X10) | ~r1_tarski(X8,X10) [superposition 72,126] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
157. ~r1_tarski(k2_xboole_0(k10_xtuple_0(X1),k10_xtuple_0(X0)),k10_xtuple_0(k1_funct_4(X1,X0))) | ~v1_relat_1(X0) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | k10_xtuple_0(k1_funct_4(X1,X0)) = k2_xboole_0(k10_xtuple_0(X1),k10_xtuple_0(X0)) | ~v1_funct_1(X0) [resolution 64,69] {a:1,w:28,wCS:11200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
170. r1_tarski(k6_subset_1(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 65,94] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
175. r1_tarski(k6_subset_1(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3) [subsumption resolution 170,99] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
176. r1_tarski(k6_subset_1(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3, 4) [subsumption resolution 175,104] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
177. r1_tarski(k6_subset_1(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) | ~v1_relat_1(sK0) <- (2, 3, 4, 5) [subsumption resolution 176,109] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
178. r1_tarski(k6_subset_1(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) <- (2, 3, 4, 5, 6) [subsumption resolution 177,114] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
180. 11 <=> r1_tarski(k6_subset_1(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) [avatar definition]
182. r1_tarski(k6_subset_1(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) <- (11) [avatar component clause 180] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
183. 11 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar split clause 178,112,107,102,97,92,180]
184. r1_tarski(k2_xboole_0(k10_xtuple_0(sK1),k10_xtuple_0(sK0)),k10_xtuple_0(k1_funct_4(sK0,sK1))) | ~r1_tarski(k10_xtuple_0(sK1),k10_xtuple_0(k1_funct_4(sK0,sK1))) <- (11) [resolution 182,144] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
186. r1_tarski(k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) | ~r1_tarski(k10_xtuple_0(sK1),k10_xtuple_0(k1_funct_4(sK0,sK1))) <- (11) [forward demodulation 184,80] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
188. 12 <=> r1_tarski(k10_xtuple_0(sK1),k10_xtuple_0(k1_funct_4(sK0,sK1))) [avatar definition]
190. ~r1_tarski(k10_xtuple_0(sK1),k10_xtuple_0(k1_funct_4(sK0,sK1))) <- (~12) [avatar component clause 188] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
192. 13 <=> r1_tarski(k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) [avatar definition]
194. r1_tarski(k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)),k10_xtuple_0(k1_funct_4(sK0,sK1))) <- (13) [avatar component clause 192] {a:2,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
195. ~12 | 13 | ~11 [avatar split clause 186,180,192,188]
206. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~12) [resolution 190,63] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
208. ~v1_relat_1(sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (5, ~12) [subsumption resolution 206,109] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
209. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (5, 6, ~12) [subsumption resolution 208,114] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
210. ~v1_relat_1(sK1) <- (3, 5, 6, ~12) [subsumption resolution 209,99] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
211. $false <- (3, 4, 5, 6, ~12) [subsumption resolution 210,104] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
212. ~3 | ~4 | ~5 | ~6 | 12 [avatar contradiction clause 211]
266. ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) = k10_xtuple_0(k1_funct_4(sK0,sK1)) | ~v1_funct_1(sK1) <- (13) [resolution 157,194] {a:3,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
275. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) = k10_xtuple_0(k1_funct_4(sK0,sK1)) | ~v1_funct_1(sK1) <- (4, 13) [subsumption resolution 266,104] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
276. ~v1_relat_1(sK0) | k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) = k10_xtuple_0(k1_funct_4(sK0,sK1)) | ~v1_funct_1(sK1) <- (4, 5, 13) [subsumption resolution 275,109] {a:3,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
277. k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) = k10_xtuple_0(k1_funct_4(sK0,sK1)) | ~v1_funct_1(sK1) <- (4, 5, 6, 13) [subsumption resolution 276,114] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
278. ~v1_funct_1(sK1) <- (~1, 4, 5, 6, 13) [subsumption resolution 277,89] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
279. $false <- (~1, 3, 4, 5, 6, 13) [subsumption resolution 278,99] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
280. 1 | ~3 | ~4 | ~5 | ~6 | ~13 [avatar contradiction clause 279]
283. $false [avatar sat refutation 90,95,100,105,110,115,183,195,212,280] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:138,thDist:-138}
% SZS output end Proof for topmetr2__t2_topmetr2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 151
% Active clauses: 44
% Passive clauses: 70
% Final active clauses: 43
% Final passive clauses: 22
% Main loop iterations started: 48
% 
% >>> Simplifying Inferences
% Duplicate literals: 42
% Fw subsumption resolutions: 20
% Fw demodulations: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Equational tautologies: 2
% Forward subsumptions: 17
% 
% >>> Generating Inferences
% Binary resolution: 35
% Unit resulting resolution: 2
% Forward superposition: 32
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 19
% SAT solver unit clauses: 6
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward demodulation: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

