i: [1,40,0,2,1,1,0,0] t184_relat_1
a: [3,67,0,2,1,40]
s: 67
i: [1,39,0,2,1,1,0,0] t184_relat_1
a: [3,72,0,2,1,39]
s: 72
i: [1,42,0,9,2,1,0,0] t184_relat_1
a: [3,57,0,3,1,42]
s: 57
a: [3,60,0,6,1,42]
s: 60
i: [1,51,0,3,1,0,0,1] reflexivity_r1_tarski
s: 51
i: [1,50,0,4,1,0,0,255] existence_m1_subset_1
s: 50
i: [1,52,0,6,2,0,0,1] dt_k5_relat_1
s: 52
i: [1,45,0,7,2,0,0,1] t59_relat_1
s: 45
i: [1,47,0,7,2,0,0,255] t3_subset
s: 47
i: [1,48,0,7,2,0,0,255] t3_subset
s: 48
i: [1,43,0,8,2,0,0,1] rd4_relat_1
s: 43
d: [2,74,1,5,1,0,40,47,50]
s: 74
i: [1,44,0,8,3,0,0,255] cc2_relat_1
s: 44
d: [2,80,1,6,2,0,40,44,50]
s: 80
i: [1,49,0,9,3,0,0,1] t1_xboole_1
s: 49
i: [1,53,0,11,2,0,0,2] rd5_relat_1
s: 53
d: [2,84,1,9,2,1,40,49,60]
s: 84
i: [1,46,0,14,4,0,0,1] t76_relat_1
s: 46
d: [2,91,2,12,3,1,40,84,49]
s: 91
d: [2,81,1,7,3,0,40,44,48]
s: 81
d: [2,83,1,10,3,0,40,49,45]
s: 83
d: [2,122,2,0,0,3,52,67,57,60,83]
e: 122
d: [2,128,2,5,2,1,40,83,60]
d: [2,139,2,2,1,2,21,128,57]
d: [2,140,2,0,0,3,21,139,67]
e: 140
r: 57
r: 91
r: 84
r: 60
a: [3,56,0,3,1,42]
s: 56
a: [3,61,0,6,1,42]
s: 61
d: [2,103,3,13,4,1,40,81,91]
a: [3,116,3,5,1,103]
s: 116
d: [2,144,4,0,0,2,52,67,116,52]
e: 144
d: [2,145,4,2,1,1,40,116,52]
d: [2,147,4,0,0,2,21,145,67]
e: 147
r: 116
d: [2,143,1,6,2,1,40,56,49]
s: 143
d: [2,95,1,15,5,0,44,46,43]
d: [2,97,1,13,4,0,18,95]
s: 97
d: [2,150,2,7,2,1,40,143,45]
d: [2,152,2,5,1,2,21,150,72]
s: 152
d: [2,158,1,11,3,0,21,97,81]
s: 158
d: [2,159,3,6,2,2,40,152,81]
d: [2,162,3,4,1,3,21,159,67]
s: 162
d: [2,93,1,15,5,0,44,46,43]
d: [2,98,1,13,4,0,18,93]
s: 98
d: [2,174,1,11,3,0,21,98,81]
s: 174
d: [2,175,2,0,0,3,52,67,56,61,174]
e: 175
d: [2,179,2,5,2,1,40,174,61]
d: [2,185,2,2,1,2,21,179,56]
d: [2,186,2,0,0,3,21,185,67]
e: 186
i: [1,41,0,9,2,1,0,0] t184_relat_1
f: 186,140,40,41,42
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t184_relat_1
% SZS output start Proof for relat_1__t184_relat_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,X1) <=> r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0)))))) [input t184_relat_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,X1) <=> r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0)))))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
9. ! [X0] : (v1_relat_1(X0) => k5_relat_1(X0,k9_xtuple_0(X0)) = X0) [input rd4_relat_1]
12. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
13. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
14. ! [X0] : ! [X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [input t59_relat_1]
15. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0))))) [input t76_relat_1]
16. ! [X0,X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [flattening 14]
17. ! [X0,X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0))))) [flattening 15]
18. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 13]
19. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 12]
21. ? [X0] : (? [X1] : ((r1_tarski(X0,X1) <~> r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0)))) & v1_relat_1(X1)) & v1_relat_1(X0)) [ennf transformation 2]
22. ! [X0] : (k5_relat_1(X0,k9_xtuple_0(X0)) = X0 | ~v1_relat_1(X0)) [ennf transformation 9]
23. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
24. ! [X0,X1] : (r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1)) [ennf transformation 16]
25. ! [X0,X1] : (! [X2] : ((r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [ennf transformation 17]
26. ! [X0,X1] : (! [X2] : (r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [flattening 25]
27. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 19]
28. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 27]
31. ? [X0] : (? [X1] : (((~r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0))) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0))) | r1_tarski(X0,X1))) & v1_relat_1(X1)) & v1_relat_1(X0)) [nnf transformation 21]
32. ? [X0] : (? [X1] : ((~r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0))) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0))) | r1_tarski(X0,X1)) & v1_relat_1(X1)) & v1_relat_1(X0)) [flattening 31]
33. ? [X0] : (? [X1] : ((~r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0))) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0))) | r1_tarski(X0,X1)) & v1_relat_1(X1)) & v1_relat_1(X0)) => (? [X1] : ((~r1_tarski(sK0,k5_relat_1(X1,k9_xtuple_0(sK0))) | ~r1_tarski(sK0,X1)) & (r1_tarski(sK0,k5_relat_1(X1,k9_xtuple_0(sK0))) | r1_tarski(sK0,X1)) & v1_relat_1(X1)) & v1_relat_1(sK0)) [choice axiom]
34. ? [X1] : ((~r1_tarski(sK0,k5_relat_1(X1,k9_xtuple_0(sK0))) | ~r1_tarski(sK0,X1)) & (r1_tarski(sK0,k5_relat_1(X1,k9_xtuple_0(sK0))) | r1_tarski(sK0,X1)) & v1_relat_1(X1)) => ((~r1_tarski(sK0,k5_relat_1(sK1,k9_xtuple_0(sK0))) | ~r1_tarski(sK0,sK1)) & (r1_tarski(sK0,k5_relat_1(sK1,k9_xtuple_0(sK0))) | r1_tarski(sK0,sK1)) & v1_relat_1(sK1)) [choice axiom]
35. ((~r1_tarski(sK0,k5_relat_1(sK1,k9_xtuple_0(sK0))) | ~r1_tarski(sK0,sK1)) & (r1_tarski(sK0,k5_relat_1(sK1,k9_xtuple_0(sK0))) | r1_tarski(sK0,sK1)) & v1_relat_1(sK1)) & v1_relat_1(sK0) [skolemisation 32,34,33]
36. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 18]
40. v1_relat_1(sK1) [cnf transformation 35] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. r1_tarski(sK0,k5_relat_1(sK1,k9_xtuple_0(sK0))) | r1_tarski(sK0,sK1) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. ~r1_tarski(sK0,k5_relat_1(sK1,k9_xtuple_0(sK0))) | ~r1_tarski(sK0,sK1) [cnf transformation 35] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
43. k5_relat_1(X0,k9_xtuple_0(X0)) = X0 | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
45. r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1) [cnf transformation 24] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
46. r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 26] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
48. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 36] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
49. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. 1 <=> r1_tarski(sK0,sK1) [avatar definition]
56. r1_tarski(sK0,sK1) <- (1) [avatar component clause 55] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
57. ~r1_tarski(sK0,sK1) <- (~1) [avatar component clause 55] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
59. 2 <=> r1_tarski(sK0,k5_relat_1(sK1,k9_xtuple_0(sK0))) [avatar definition]
60. r1_tarski(sK0,k5_relat_1(sK1,k9_xtuple_0(sK0))) <- (2) [avatar component clause 59] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
61. ~r1_tarski(sK0,k5_relat_1(sK1,k9_xtuple_0(sK0))) <- (~2) [avatar component clause 59] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
62. ~1 | ~2 [avatar split clause 42,59,55]
63. 1 | 2 [avatar split clause 41,59,55]
65. 3 <=> v1_relat_1(sK1) [avatar definition]
67. v1_relat_1(sK1) <- (3) [avatar component clause 65] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
68. 3 [avatar split clause 40,65]
81. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 44,48] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. ~r1_tarski(X2,k5_relat_1(X3,X4)) | r1_tarski(X2,X3) | ~v1_relat_1(X3) [resolution 49,45] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
93. r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0))) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) | ~v1_relat_1(X0) [superposition 46,43] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
98. r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0))) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [duplicate literal removal 93] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
128. r1_tarski(sK0,sK1) | ~v1_relat_1(sK1) <- (2) [resolution 83,60] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
139. ~v1_relat_1(sK1) <- (~1, 2) [subsumption resolution 128,57] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
140. $false <- (~1, 2, 3) [subsumption resolution 139,67] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
141. 1 | ~2 | ~3 [avatar contradiction clause 140]
174. r1_tarski(X0,k5_relat_1(X1,k9_xtuple_0(X0))) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) [subsumption resolution 98,81] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
179. ~r1_tarski(sK0,sK1) | ~v1_relat_1(sK1) <- (~2) [resolution 174,61] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
185. ~v1_relat_1(sK1) <- (1, ~2) [subsumption resolution 179,56] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
186. $false <- (1, ~2, 3) [subsumption resolution 185,67] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
187. ~1 | 2 | ~3 [avatar contradiction clause 186]
188. $false [avatar sat refutation 62,63,68,141,187] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:97,thDist:-97}
% SZS output end Proof for relat_1__t184_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 118
% Active clauses: 29
% Passive clauses: 46
% Final active clauses: 24
% Final passive clauses: 14
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Duplicate literals: 12
% Fw subsumption resolutions: 17
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 41
% Unit resulting resolution: 3
% Forward superposition: 13
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 2
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption index maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

