i: [1,180,0,2,1,1,0,0] t12_xxreal_1
a: [3,290,0,2,1,180]
s: 290
i: [1,179,0,2,1,1,0,0] t12_xxreal_1
a: [3,295,0,2,1,179]
s: 295
i: [1,182,0,8,2,1,0,0] t12_xxreal_1
a: [3,276,0,5,1,182]
s: 276
i: [1,181,0,5,1,1,0,0] t12_xxreal_1
a: [3,285,0,5,1,181]
s: 285
i: [1,183,0,2,1,0,0,3] fc5_numbers
a: [3,300,0,2,1,183]
s: 300
i: [1,184,0,2,1,0,0,2] fc2_membered
a: [3,305,0,2,1,184]
s: 305
i: [1,185,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,310,0,2,1,185]
s: 310
i: [1,186,0,2,1,0,0,255] fc1_xboole_0
a: [3,315,0,2,1,186]
s: 315
i: [1,232,0,7,3,0,0,2] reflexivity_r1_xxreal_0
a: [3,324,0,2,1,232]
s: 324
d: [2,456,1,0,0,2,40,324,295]
e: 456
d: [2,455,1,0,0,2,40,324,290]
e: 455
d: [2,454,1,0,0,2,52,290,324]
e: 454
r: 324
i: [1,246,0,2,1,0,0,2] rc2_xboole_0
a: [3,332,0,2,1,246]
s: 332
i: [1,247,0,2,1,0,0,255] rc1_xreal_0
a: [3,337,0,2,1,247]
s: 337
i: [1,248,0,2,1,0,0,2] rc1_xboole_0
a: [3,342,0,2,1,248]
s: 342
i: [1,249,0,2,1,0,0,1] rc1_xxreal_0
a: [3,347,0,2,1,249]
s: 347
i: [1,251,0,2,1,0,0,255] rc1_membered
a: [3,352,0,2,1,251]
s: 352
i: [1,250,0,2,1,0,0,255] rc1_membered
a: [3,357,0,2,1,250]
s: 357
i: [1,253,0,2,1,0,0,255] rc2_membered
a: [3,362,0,2,1,253]
s: 362
i: [1,252,0,2,1,0,0,255] rc2_membered
a: [3,367,0,2,1,252]
s: 367
i: [1,255,0,2,1,0,0,255] rc3_membered
a: [3,372,0,2,1,255]
s: 372
i: [1,254,0,2,1,0,0,255] rc3_membered
a: [3,377,0,2,1,254]
s: 377
i: [1,258,0,2,1,0,0,255] rc2_xreal_0
a: [3,382,0,2,1,258]
s: 382
i: [1,257,0,2,1,0,0,255] rc2_xreal_0
a: [3,387,0,2,1,257]
s: 387
i: [1,256,0,2,1,0,0,255] rc2_xreal_0
a: [3,392,0,2,1,256]
s: 392
i: [1,261,0,2,1,0,0,255] rc3_xreal_0
a: [3,397,0,2,1,261]
s: 397
i: [1,260,0,2,1,0,0,255] rc3_xreal_0
a: [3,402,0,2,1,260]
s: 402
i: [1,259,0,2,1,0,0,255] rc3_xreal_0
a: [3,407,0,2,1,259]
s: 407
i: [1,264,0,2,1,0,0,255] rc4_xreal_0
a: [3,412,0,2,1,264]
s: 412
i: [1,263,0,2,1,0,0,255] rc4_xreal_0
a: [3,417,0,2,1,263]
s: 417
i: [1,262,0,2,1,0,0,255] rc4_xreal_0
a: [3,422,0,2,1,262]
s: 422
i: [1,266,0,2,1,0,0,2] rc4_xxreal_0
a: [3,427,0,2,1,266]
s: 427
i: [1,265,0,2,1,0,0,2] rc4_xxreal_0
a: [3,432,0,2,1,265]
s: 432
i: [1,268,0,2,1,0,0,255] rc3_xxreal_0
a: [3,437,0,2,1,268]
s: 437
i: [1,267,0,2,1,0,0,255] rc3_xxreal_0
a: [3,442,0,2,1,267]
s: 442
i: [1,270,0,2,1,0,0,255] rc2_xxreal_0
a: [3,447,0,2,1,270]
s: 447
i: [1,269,0,2,1,0,0,255] rc2_xxreal_0
a: [3,452,0,2,1,269]
s: 452
i: [1,187,0,3,1,0,0,255] d2_xboole_0
a: [3,320,0,3,1,187]
s: 320
r: 310
i: [1,188,0,4,2,0,0,255] cc2_membered
s: 188
i: [1,189,0,4,2,0,0,255] cc3_membered
s: 189
i: [1,190,0,4,2,0,0,255] cc2_xreal_0
s: 190
i: [1,191,0,4,2,0,0,255] cc2_xxreal_0
s: 191
i: [1,192,0,4,2,0,0,255] cc4_membered
s: 192
i: [1,195,0,4,2,0,0,255] cc1_membered
s: 195
d: [2,461,1,2,1,1,40,195,352]
a: [3,467,1,2,1,461]
s: 467
i: [1,197,0,4,2,0,0,255] cc4_xreal_0
s: 197
d: [2,462,1,2,1,1,40,195,362]
a: [3,472,1,2,1,462]
s: 472
i: [1,198,0,4,2,0,0,255] cc18_membered
s: 198
i: [1,230,0,4,1,0,0,2] existence_m1_subset_1
s: 230
d: [2,463,1,2,1,1,40,195,372]
a: [3,477,1,2,1,463]
s: 477
i: [1,199,0,5,2,0,0,255] t6_boole
s: 199
d: [2,485,1,2,1,1,40,197,337]
a: [3,492,1,2,1,485]
s: 492
i: [1,208,0,5,2,0,0,3] cc1_xxreal_0
s: 208
i: [1,227,0,5,2,0,0,1] t7_boole
s: 227
d: [2,538,1,4,1,1,40,227,285]
a: [3,542,1,4,1,538]
s: 542
a: [3,327,0,5,2,232]
s: 327
d: [2,500,1,2,1,1,40,198,315]
a: [3,507,1,2,1,500]
s: 507
i: [1,214,0,6,3,0,0,255] cc7_xxreal_0
s: 214
i: [1,215,0,6,3,0,0,255] cc7_xxreal_0
s: 215
d: [2,501,1,2,1,1,40,198,342]
a: [3,512,1,2,1,501]
s: 512
i: [1,218,0,6,3,0,0,255] cc3_xxreal_0
s: 218
d: [2,502,1,2,1,1,40,198,422]
a: [3,517,1,2,1,502]
s: 517
i: [1,220,0,6,3,0,0,255] cc3_xxreal_0
s: 220
i: [1,221,0,6,3,0,0,255] cc5_xxreal_0
s: 221
d: [2,503,1,2,1,1,40,198,432]
a: [3,522,1,2,1,503]
s: 522
i: [1,223,0,6,3,0,0,255] cc5_xxreal_0
s: 223
d: [2,551,1,4,2,1,40,214,387]
d: [2,553,1,2,1,2,21,551,392]
a: [3,557,1,2,1,553]
s: 557
i: [1,225,0,6,2,0,0,1] t1_subset
s: 225
i: [1,231,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 231
d: [2,611,1,5,1,1,40,225,285]
a: [3,615,1,5,1,611]
s: 615
i: [1,193,0,7,3,0,0,255] cc14_membered
s: 193
d: [2,624,2,6,2,1,40,193,615]
a: [3,628,2,4,1,624]
s: 628
i: [1,194,0,7,3,0,0,2] cc13_membered
s: 194
i: [1,196,0,7,3,0,0,255] cc17_membered
s: 196
d: [2,635,2,6,2,1,40,194,615]
a: [3,639,2,4,1,635]
s: 639
i: [1,226,0,7,3,0,0,1] t8_boole
s: 226
d: [2,646,2,6,2,1,40,196,615]
a: [3,650,2,4,1,646]
s: 650
i: [1,228,0,7,2,0,0,1] t2_tarski
s: 228
i: [1,210,0,8,4,0,0,255] cc4_xxreal_0
s: 210
d: [2,659,1,6,3,1,40,210,295]
a: [3,672,1,2,1,659]
s: 672
r: 672
i: [1,212,0,8,4,0,0,255] cc6_xxreal_0
s: 212
a: [3,671,1,2,1,659]
s: 671
i: [1,216,0,8,4,0,0,255] cc8_xxreal_0
s: 216
i: [1,224,0,8,3,0,0,1] t2_subset
s: 224
a: [3,676,1,2,1,659]
s: 676
i: [1,233,0,8,3,0,0,1] fc3_xxreal_1
s: 233
d: [2,730,3,0,0,3,52,295,290,639,233]
e: 730
d: [2,731,3,4,2,1,40,233,639]
d: [2,733,3,2,1,2,21,731,290]
d: [2,734,3,0,0,3,21,733,295]
e: 734
r: 639
d: [2,660,1,6,3,1,40,210,290]
a: [3,685,1,2,1,660]
s: 685
r: 685
i: [1,234,0,8,3,0,0,1] fc1_xxreal_1
s: 234
i: [1,229,0,10,3,0,0,1] t2_tarski
s: 229
a: [3,684,1,2,1,660]
s: 684
i: [1,235,0,10,4,0,0,2] connectedness_r1_xxreal_0
s: 235
a: [3,689,1,2,1,660]
s: 689
i: [1,200,0,11,3,0,0,255] d1_xxreal_1
s: 200
i: [1,201,0,11,3,0,0,255] d3_xxreal_1
s: 201
a: [3,643,2,2,1,635]
s: 643
i: [1,202,0,14,5,0,0,1] t2_xxreal_1
s: 202
d: [2,746,3,6,3,1,40,643,210]
a: [3,750,3,2,1,746]
s: 750
r: 750
i: [1,203,0,14,5,0,0,1] t2_xxreal_1
s: 203
i: [1,205,0,14,5,0,0,1] t1_xxreal_1
s: 205
a: [3,749,3,2,1,746]
s: 749
i: [1,206,0,14,5,0,0,1] t1_xxreal_1
s: 206
a: [3,754,3,2,1,746]
s: 754
i: [1,236,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 236
i: [1,237,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 237
d: [2,760,1,9,4,1,40,202,285]
d: [2,762,1,7,3,2,21,760,290]
d: [2,763,1,5,2,3,21,762,295]
d: [2,764,1,3,1,4,21,763,643]
a: [3,768,1,3,1,764]
s: 768
i: [1,238,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 238
d: [2,776,1,9,4,1,40,203,285]
d: [2,778,1,7,3,2,21,776,290]
d: [2,779,1,5,2,3,21,778,295]
d: [2,780,1,3,1,4,21,779,643]
a: [3,784,1,3,1,780]
s: 784
i: [1,239,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 239
i: [1,241,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 241
d: [2,814,2,7,3,1,40,768,235]
d: [2,815,2,5,2,2,21,814,643]
d: [2,816,2,3,1,3,21,815,295]
a: [3,820,2,3,1,816]
s: 820
i: [1,242,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 242
a: [3,638,2,4,1,635]
s: 638
i: [1,243,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 243
i: [1,244,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 244
d: [2,617,1,5,1,1,40,231,285]
a: [3,621,1,5,1,617]
s: 621
i: [1,204,0,17,6,0,0,1] t2_xxreal_1
s: 204
d: [2,552,1,4,2,1,40,214,447]
d: [2,559,1,2,1,2,21,552,452]
a: [3,563,1,2,1,559]
s: 563
i: [1,207,0,17,6,0,0,1] t1_xxreal_1
s: 207
d: [2,858,3,0,0,6,52,643,295,290,820,784,276,207]
e: 858
d: [2,859,1,12,5,1,40,207,276]
d: [2,868,1,9,4,2,21,859,784]
d: [2,869,1,6,3,3,21,868,820]
d: [2,870,1,4,2,4,21,869,290]
d: [2,871,1,2,1,5,21,870,295]
d: [2,872,1,0,0,6,21,871,643]
e: 872
f: 872,816,780,764,734,635,611,179,180,181,182
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_1__t12_xxreal_1
% SZS output start Proof for xxreal_1__t12_xxreal_1
1. ! [X0] : ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k3_xxreal_1(X1,X2)) => (X0 != X1 & r2_hidden(X0,k1_xxreal_1(X1,X2)))))) [input t12_xxreal_1]
2. ~! [X0] : ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k3_xxreal_1(X1,X2)) => (X0 != X1 & r2_hidden(X0,k1_xxreal_1(X1,X2)))))) [negated conjecture 1]
5. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xxreal_0(X1))) [input cc13_membered]
28. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
42. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v2_membered(k3_xxreal_1(X0,X1))) [input fc3_xxreal_1]
60. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
61. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k1_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0)))))) [input t1_xxreal_1]
64. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k3_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1)))))) [input t2_xxreal_1]
68. ~! [X0,X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k3_xxreal_1(X1,X2)) => (X0 != X1 & r2_hidden(X0,k1_xxreal_1(X1,X2)))))) [flattening 2]
70. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 60]
84. ? [X0,X1] : (? [X2] : (((X0 = X1 | ~r2_hidden(X0,k1_xxreal_1(X1,X2))) & r2_hidden(X0,k3_xxreal_1(X1,X2))) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) [ennf transformation 68]
85. ? [X0,X1] : (? [X2] : ((X0 = X1 | ~r2_hidden(X0,k1_xxreal_1(X1,X2))) & r2_hidden(X0,k3_xxreal_1(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) [flattening 84]
92. ! [X0] : (! [X1] : (v1_xxreal_0(X1) | ~m1_subset_1(X1,X0)) | ~v2_membered(X0)) [ennf transformation 5]
100. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X0,k3_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 64]
101. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X0,k1_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 61]
117. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 70]
124. ! [X0,X1] : (v2_membered(k3_xxreal_1(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 42]
125. ! [X0,X1] : (v2_membered(k3_xxreal_1(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 124]
128. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 28]
129. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 128]
134. ? [X0,X1] : (? [X2] : ((X0 = X1 | ~r2_hidden(X0,k1_xxreal_1(X1,X2))) & r2_hidden(X0,k3_xxreal_1(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : ((sK0 = sK1 | ~r2_hidden(sK0,k1_xxreal_1(sK1,X2))) & r2_hidden(sK0,k3_xxreal_1(sK1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
135. ? [X2] : ((sK0 = sK1 | ~r2_hidden(sK0,k1_xxreal_1(sK1,X2))) & r2_hidden(sK0,k3_xxreal_1(sK1,X2)) & v1_xxreal_0(X2)) => ((sK0 = sK1 | ~r2_hidden(sK0,k1_xxreal_1(sK1,sK2))) & r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) & v1_xxreal_0(sK2)) [choice axiom]
136. ((sK0 = sK1 | ~r2_hidden(sK0,k1_xxreal_1(sK1,sK2))) & r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1) [skolemisation 85,135,134]
137. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k3_xxreal_1(X1,X2)) | (~r1_xxreal_0(X0,X2) | r1_xxreal_0(X0,X1))) & ((r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1)) | ~r2_hidden(X0,k3_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 100]
138. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k3_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X2) | r1_xxreal_0(X0,X1)) & ((r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1)) | ~r2_hidden(X0,k3_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 137]
139. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k1_xxreal_1(X1,X2)) | (~r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X0))) & ((r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0)) | ~r2_hidden(X0,k1_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 101]
140. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k1_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X0)) & ((r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0)) | ~r2_hidden(X0,k1_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 139]
179. v1_xxreal_0(sK1) [cnf transformation 136] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
180. v1_xxreal_0(sK2) [cnf transformation 136] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
181. r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) [cnf transformation 136] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
182. sK0 = sK1 | ~r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) [cnf transformation 136] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
194. ~m1_subset_1(X1,X0) | v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 92] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
202. ~r2_hidden(X0,k3_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 138] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
203. ~r2_hidden(X0,k3_xxreal_1(X1,X2)) | r1_xxreal_0(X0,X2) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 138] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
207. r2_hidden(X0,k1_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 140] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
225. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 117] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
233. v2_membered(k3_xxreal_1(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 125] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
235. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 129] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
274. 1 <=> r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) [avatar definition]
276. ~r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) <- (~1) [avatar component clause 274] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
278. 2 <=> sK0 = sK1 [avatar definition]
281. ~1 | 2 [avatar split clause 182,278,274]
283. 3 <=> r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) [avatar definition]
285. r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) <- (3) [avatar component clause 283] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
286. 3 [avatar split clause 181,283]
288. 4 <=> v1_xxreal_0(sK2) [avatar definition]
290. v1_xxreal_0(sK2) <- (4) [avatar component clause 288] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
291. 4 [avatar split clause 180,288]
293. 5 <=> v1_xxreal_0(sK1) [avatar definition]
295. v1_xxreal_0(sK1) <- (5) [avatar component clause 293] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
296. 5 [avatar split clause 179,293]
611. m1_subset_1(sK0,k3_xxreal_1(sK1,sK2)) <- (3) [resolution 225,285] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
613. 61 <=> m1_subset_1(sK0,k3_xxreal_1(sK1,sK2)) [avatar definition]
615. m1_subset_1(sK0,k3_xxreal_1(sK1,sK2)) <- (61) [avatar component clause 613] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
616. 61 | ~3 [avatar split clause 611,283,613]
635. v1_xxreal_0(sK0) | ~v2_membered(k3_xxreal_1(sK1,sK2)) <- (61) [resolution 194,615] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
637. 65 <=> v2_membered(k3_xxreal_1(sK1,sK2)) [avatar definition]
639. ~v2_membered(k3_xxreal_1(sK1,sK2)) <- (~65) [avatar component clause 637] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
641. 66 <=> v1_xxreal_0(sK0) [avatar definition]
643. v1_xxreal_0(sK0) <- (66) [avatar component clause 641] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
644. ~65 | 66 | ~61 [avatar split clause 635,613,641,637]
731. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~65) [resolution 233,639] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
733. ~v1_xxreal_0(sK1) <- (4, ~65) [subsumption resolution 731,290] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
734. $false <- (4, 5, ~65) [subsumption resolution 733,295] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
735. ~4 | ~5 | 65 [avatar contradiction clause 734]
760. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3) [resolution 202,285] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
762. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3, 4) [subsumption resolution 760,290] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
763. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK0) <- (3, 4, 5) [subsumption resolution 762,295] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
764. ~r1_xxreal_0(sK0,sK1) <- (3, 4, 5, 66) [subsumption resolution 763,643] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
766. 86 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
768. ~r1_xxreal_0(sK0,sK1) <- (~86) [avatar component clause 766] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
769. ~86 | ~3 | ~4 | ~5 | ~66 [avatar split clause 764,641,293,288,283,766]
776. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3) [resolution 203,285] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
778. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3, 4) [subsumption resolution 776,290] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
779. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK0) <- (3, 4, 5) [subsumption resolution 778,295] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
780. r1_xxreal_0(sK0,sK2) <- (3, 4, 5, 66) [subsumption resolution 779,643] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
782. 88 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
784. r1_xxreal_0(sK0,sK2) <- (88) [avatar component clause 782] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
785. 88 | ~3 | ~4 | ~5 | ~66 [avatar split clause 780,641,293,288,283,782]
814. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~86) [resolution 768,235] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
815. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) <- (66, ~86) [subsumption resolution 814,643] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
816. r1_xxreal_0(sK1,sK0) <- (5, 66, ~86) [subsumption resolution 815,295] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
818. 93 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
820. r1_xxreal_0(sK1,sK0) <- (93) [avatar component clause 818] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
821. 93 | ~5 | ~66 | 86 [avatar split clause 816,766,641,293,818]
859. ~r1_xxreal_0(sK0,sK2) | ~r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~1) [resolution 207,276] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
868. ~r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~1, 88) [subsumption resolution 859,784] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
869. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~1, 88, 93) [subsumption resolution 868,820] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
870. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~1, 4, 88, 93) [subsumption resolution 869,290] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
871. ~v1_xxreal_0(sK0) <- (~1, 4, 5, 88, 93) [subsumption resolution 870,295] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
872. $false <- (~1, 4, 5, 66, 88, 93) [subsumption resolution 871,643] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
873. 1 | ~4 | ~5 | ~66 | ~88 | ~93 [avatar contradiction clause 872]
874. sK0 != sK1 | r1_xxreal_0(sK0,sK1) | ~r1_xxreal_0(sK1,sK0) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
875. $false [avatar sat refutation 281,286,291,296,616,644,735,769,785,821,873,874] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:437,thDist:-437}
% SZS output end Proof for xxreal_1__t12_xxreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 66
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 5
% 
% >>> Saturation
% Initial clauses: 92
% Generated clauses: 316
% Active clauses: 110
% Passive clauses: 162
% Final active clauses: 104
% Final passive clauses: 47
% Main loop iterations started: 115
% 
% >>> Simplifying Inferences
% Duplicate literals: 36
% Fw subsumption resolutions: 25
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Forward subsumptions: 22
% 
% >>> Generating Inferences
% Binary resolution: 90
% Unit resulting resolution: 3
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 14
% Split components: 33
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 88
% SAT solver unit clauses: 34
% SAT solver binary clauses: 30
% 
% Memory used [KB]: 5500
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% superposition: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

