i: [1,28,0,10,1,1,0,0] t44_xtuple_0
a: [3,40,0,10,1,28]
s: 40
i: [1,34,0,3,1,0,0,1] reflexivity_r1_tarski
s: 34
i: [1,31,0,5,1,0,0,1] t17_xboole_1
s: 31
i: [1,35,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 35
i: [1,29,0,6,1,0,0,255] d14_xtuple_0
s: 29
i: [1,30,0,6,1,0,0,255] d17_xtuple_0
s: 30
i: [1,36,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 36
i: [1,32,0,8,2,0,0,1] t17_xtuple_0
s: 32
d: [2,44,1,5,1,0,44,31,36]
s: 44
i: [1,33,0,11,3,0,0,1] t19_xboole_1
s: 33
d: [2,49,1,14,2,1,40,33,40]
a: [3,57,1,7,1,49]
s: 57
d: [2,63,2,0,0,1,52,31,57,32]
e: 63
d: [2,64,2,5,1,1,40,57,32]
d: [2,66,2,0,0,1,21,64,31]
e: 66
r: 57
a: [3,56,1,7,1,49]
s: 56
a: [3,61,1,7,1,49]
s: 61
d: [2,68,2,0,0,1,52,44,61,32]
e: 68
d: [2,69,2,5,1,1,40,61,32]
d: [2,71,2,0,0,1,21,69,44]
e: 71
f: 71,66,49,28
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t44_xtuple_0
% SZS output start Proof for xtuple_0__t44_xtuple_0
1. ! [X0] : ! [X1] : r1_tarski(k14_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k14_xtuple_0(X0),k14_xtuple_0(X1))) [input t44_xtuple_0]
2. ~! [X0] : ! [X1] : r1_tarski(k14_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k14_xtuple_0(X0),k14_xtuple_0(X1))) [negated conjecture 1]
3. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
13. ! [X0] : ! [X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [input t17_xboole_1]
14. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1))) [input t17_xtuple_0]
15. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k3_xboole_0(X1,X2))) [input t19_xboole_1]
16. ~! [X0,X1] : r1_tarski(k14_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k14_xtuple_0(X0),k14_xtuple_0(X1))) [flattening 2]
17. ! [X0,X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [flattening 13]
18. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1))) [flattening 14]
19. ! [X0,X1,X2] : ((r1_tarski(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k3_xboole_0(X1,X2))) [flattening 15]
22. ? [X0,X1] : ~r1_tarski(k14_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k14_xtuple_0(X0),k14_xtuple_0(X1))) [ennf transformation 16]
23. ! [X0,X1] : (r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 18]
24. ! [X0,X1,X2] : (r1_tarski(X0,k3_xboole_0(X1,X2)) | (~r1_tarski(X0,X2) | ~r1_tarski(X0,X1))) [ennf transformation 19]
25. ! [X0,X1,X2] : (r1_tarski(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X2) | ~r1_tarski(X0,X1)) [flattening 24]
26. ? [X0,X1] : ~r1_tarski(k14_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k14_xtuple_0(X0),k14_xtuple_0(X1))) => ~r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k14_xtuple_0(sK0),k14_xtuple_0(sK1))) [choice axiom]
27. ~r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k14_xtuple_0(sK0),k14_xtuple_0(sK1))) [skolemisation 22,26]
28. ~r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k14_xtuple_0(sK0),k14_xtuple_0(sK1))) [cnf transformation 27] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. r1_tarski(k3_xboole_0(X0,X1),X0) [cnf transformation 17] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
32. r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 23] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. r1_tarski(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 25] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
36. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. 1 <=> r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k14_xtuple_0(sK0),k14_xtuple_0(sK1))) [avatar definition]
40. ~r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k14_xtuple_0(sK0),k14_xtuple_0(sK1))) <- (~1) [avatar component clause 38] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
41. ~1 [avatar split clause 28,38]
44. r1_tarski(k3_xboole_0(X1,X0),X0) [superposition 31,36] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
49. ~r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k14_xtuple_0(sK1)) | ~r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k14_xtuple_0(sK0)) <- (~1) [resolution 33,40] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
55. 2 <=> r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k14_xtuple_0(sK0)) [avatar definition]
57. ~r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k14_xtuple_0(sK0)) <- (~2) [avatar component clause 55] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
59. 3 <=> r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k14_xtuple_0(sK1)) [avatar definition]
61. ~r1_tarski(k14_xtuple_0(k3_xboole_0(sK0,sK1)),k14_xtuple_0(sK1)) <- (~3) [avatar component clause 59] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
62. ~2 | ~3 | 1 [avatar split clause 49,38,59,55]
64. ~r1_tarski(k3_xboole_0(sK0,sK1),sK0) <- (~2) [resolution 57,32] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66. $false <- (~2) [subsumption resolution 64,31] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
67. 2 [avatar contradiction clause 66]
69. ~r1_tarski(k3_xboole_0(sK0,sK1),sK1) <- (~3) [resolution 61,32] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
71. $false <- (~3) [subsumption resolution 69,44] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
72. 3 [avatar contradiction clause 71]
73. $false [avatar sat refutation 41,62,67,72] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:28,thDist:-28}
% SZS output end Proof for xtuple_0__t44_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 9
% Generated clauses: 31
% Active clauses: 13
% Passive clauses: 15
% Final active clauses: 12
% Final passive clauses: 2
% Main loop iterations started: 13
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 2
% Forward superposition: 6
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 8
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

