i: [1,122,0,2,1,1,0,0] t4_algseq_1
a: [3,190,0,2,1,122]
s: 190
i: [1,121,0,2,1,1,0,0] t4_algseq_1
a: [3,195,0,2,1,121]
s: 195
i: [1,124,0,8,2,1,0,0] t4_algseq_1
a: [3,180,0,3,1,124]
s: 180
a: [3,183,0,5,1,124]
s: 183
i: [1,125,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,200,0,2,1,125]
s: 200
i: [1,129,0,2,1,0,0,2] fc6_ordinal1
a: [3,215,0,2,1,129]
s: 215
i: [1,128,0,2,1,0,0,2] fc6_ordinal1
a: [3,220,0,2,1,128]
s: 220
i: [1,159,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,224,0,2,1,159]
s: 224
i: [1,161,0,7,3,0,0,2] reflexivity_r1_ordinal1
a: [3,231,0,2,1,161]
s: 231
d: [2,297,1,0,0,2,40,231,215]
e: 297
d: [2,296,1,0,0,2,52,215,231]
e: 296
r: 231
i: [1,165,0,2,1,0,0,3] rc1_xreal_0
a: [3,239,0,2,1,165]
s: 239
i: [1,166,0,2,1,0,0,1] rc4_ordinal1
a: [3,244,0,2,1,166]
s: 244
i: [1,167,0,2,1,0,0,2] rc1_ordinal1
a: [3,249,0,2,1,167]
s: 249
i: [1,171,0,2,1,0,0,255] rc2_ordinal1
a: [3,254,0,2,1,171]
s: 254
i: [1,170,0,2,1,0,0,255] rc2_ordinal1
a: [3,259,0,2,1,170]
s: 259
i: [1,169,0,2,1,0,0,255] rc2_ordinal1
a: [3,264,0,2,1,169]
s: 264
i: [1,168,0,2,1,0,0,255] rc2_ordinal1
a: [3,269,0,2,1,168]
s: 269
i: [1,173,0,2,1,0,0,1] rc5_ordinal1
a: [3,274,0,2,1,173]
s: 274
i: [1,172,0,2,1,0,0,1] rc5_ordinal1
a: [3,279,0,2,1,172]
s: 279
i: [1,176,0,2,1,0,0,255] rc4_xreal_0
a: [3,284,0,2,1,176]
s: 284
i: [1,175,0,2,1,0,0,255] rc4_xreal_0
a: [3,289,0,2,1,175]
s: 289
d: [2,300,0,0,0,2,21,289,224]
e: 300
r: 224
s: 289
i: [1,174,0,2,1,0,0,255] rc4_xreal_0
a: [3,294,0,2,1,174]
s: 294
i: [1,157,0,3,1,0,0,1] reflexivity_r1_tarski
s: 157
i: [1,126,0,3,1,0,0,255] d2_xboole_0
a: [3,205,0,3,1,126]
s: 205
i: [1,127,0,3,1,0,0,255] t1_numerals
a: [3,210,0,3,1,127]
s: 210
d: [2,302,0,3,1,2,22,210,205]
a: [3,306,0,3,1,302]
s: 306
i: [1,131,0,4,2,0,0,2] cc4_xreal_0
s: 131
i: [1,132,0,4,2,0,0,255] cc9_ordinal1
s: 132
d: [2,308,1,2,1,1,40,131,239]
a: [3,313,1,2,1,308]
s: 313
i: [1,133,0,4,2,0,0,1] cc7_ordinal1
s: 133
i: [1,134,0,4,2,0,0,2] cc3_ordinal1
s: 134
d: [2,315,1,2,1,1,40,132,200]
a: [3,320,1,2,1,315]
s: 320
i: [1,136,0,4,2,0,0,3] cc2_xreal_0
s: 136
d: [2,351,1,2,1,1,40,136,195]
a: [3,358,1,2,1,351]
s: 358
i: [1,137,0,4,2,0,0,1] cc6_ordinal1
s: 137
i: [1,143,0,4,2,0,0,255] cc1_ordinal1
s: 143
d: [2,352,1,2,1,1,40,136,190]
a: [3,363,1,2,1,352]
s: 363
i: [1,144,0,4,2,0,0,255] cc1_ordinal1
s: 144
d: [2,381,1,2,1,1,40,137,195]
a: [3,388,1,2,1,381]
s: 388
i: [1,156,0,4,1,0,0,2] existence_m1_subset_1
s: 156
i: [1,135,0,5,2,0,0,255] t6_boole
s: 135
d: [2,382,1,2,1,1,40,137,190]
a: [3,393,1,2,1,382]
s: 393
i: [1,146,0,5,2,0,0,2] cc8_ordinal1
s: 146
d: [2,375,2,2,1,1,40,358,131]
a: [3,379,2,2,1,375]
s: 379
i: [1,153,0,5,2,0,0,3] t7_boole
s: 153
a: [3,227,0,5,2,159]
s: 227
d: [2,418,2,2,1,1,40,363,131]
a: [3,422,2,2,1,418]
s: 422
a: [3,234,0,5,2,161]
s: 234
d: [2,437,2,2,1,1,40,388,144]
a: [3,442,2,2,1,437]
s: 442
d: [2,449,0,5,2,1,22,135,205]
s: 449
i: [1,138,0,6,2,0,0,255] d6_card_1
s: 138
d: [2,438,2,2,1,1,40,388,143]
a: [3,447,2,2,1,438]
s: 447
i: [1,147,0,6,3,0,0,255] cc2_ordinal1
s: 147
d: [2,450,2,2,1,1,40,393,144]
a: [3,455,2,2,1,450]
s: 455
i: [1,148,0,6,2,0,0,3] t1_subset
s: 148
i: [1,158,0,6,2,0,0,3] antisymmetry_r2_hidden
s: 158
d: [2,451,2,2,1,1,40,393,143]
a: [3,460,2,2,1,451]
s: 460
i: [1,139,0,7,2,0,0,255] redefinition_k7_card_1
s: 139
d: [2,316,1,2,1,1,40,132,294]
a: [3,325,1,2,1,316]
s: 325
i: [1,140,0,7,2,0,0,1] dt_k7_card_1
s: 140
i: [1,145,0,7,3,0,0,2] cc5_ordinal1
s: 145
d: [2,327,1,2,1,1,40,133,200]
a: [3,332,1,2,1,327]
s: 332
i: [1,150,0,7,2,0,0,1] t3_subset
s: 150
d: [2,328,1,2,1,1,40,133,294]
a: [3,337,1,2,1,328]
s: 337
i: [1,151,0,7,2,0,0,1] t3_subset
s: 151
i: [1,152,0,7,3,0,0,255] t8_boole
s: 152
d: [2,339,1,2,1,1,40,134,200]
a: [3,344,1,2,1,339]
s: 344
i: [1,130,0,8,3,0,0,255] cc10_ordinal1
s: 130
d: [2,340,1,2,1,1,40,134,294]
a: [3,349,1,2,1,340]
s: 349
i: [1,149,0,8,3,0,0,3] t2_subset
s: 149
i: [1,155,0,9,3,0,0,2] t5_subset
s: 155
d: [2,353,1,2,1,1,40,136,244]
a: [3,368,1,2,1,353]
s: 368
i: [1,141,0,10,4,0,0,1] t39_nat_1
s: 141
d: [2,354,1,2,1,1,40,136,274]
a: [3,373,1,2,1,354]
s: 373
i: [1,142,0,10,4,0,0,1] t39_nat_1
s: 142
i: [1,154,0,10,3,0,0,2] t4_subset
s: 154
d: [2,383,1,2,1,1,40,137,244]
a: [3,398,1,2,1,383]
s: 398
i: [1,160,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 160
r: 180
r: 183
a: [3,179,0,3,1,124]
s: 179
r: 179
s: 180
s: 183
d: [2,599,1,7,3,1,40,160,180]
d: [2,603,1,5,2,2,21,599,379]
d: [2,604,1,3,1,3,21,603,422]
a: [3,608,1,3,1,604]
s: 608
r: 180
r: 183
r: 608
s: 179
a: [3,607,1,3,1,604]
s: 607
a: [3,184,0,5,1,124]
s: 184
i: [1,162,0,10,4,0,0,2] connectedness_r1_ordinal1
s: 162
d: [2,610,1,7,3,1,40,179,141]
d: [2,611,1,5,2,2,21,610,190]
d: [2,612,1,3,1,3,21,611,195]
a: [3,616,1,3,1,612]
s: 616
i: [1,163,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 163
r: 179
r: 184
r: 607
r: 616
s: 180
s: 183
i: [1,164,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 164
s: 608
a: [3,615,1,3,1,612]
s: 615
d: [2,658,1,11,3,1,40,164,183]
a: [3,663,1,3,1,658]
s: 663
d: [2,384,1,2,1,1,40,137,274]
a: [3,403,1,2,1,384]
s: 403
d: [2,405,1,2,1,1,40,143,215]
a: [3,411,1,2,1,405]
s: 411
d: [2,622,2,7,3,1,40,608,141]
d: [2,623,2,5,2,2,21,622,195]
d: [2,624,2,3,1,3,21,623,190]
a: [3,628,2,3,1,624]
s: 628
r: 180
r: 183
r: 608
r: 615
r: 628
s: 179
s: 607
s: 184
s: 616
a: [3,627,2,3,1,624]
s: 627
d: [2,406,1,2,1,1,40,143,249]
a: [3,416,1,2,1,406]
s: 416
d: [2,643,2,7,3,1,40,163,616]
d: [2,646,2,5,2,2,21,643,393]
d: [2,647,2,3,1,3,21,646,388]
a: [3,651,2,3,1,647]
s: 651
d: [2,424,1,2,1,1,40,144,215]
a: [3,430,1,2,1,424]
s: 430
d: [2,425,1,2,1,1,40,144,249]
a: [3,435,1,2,1,425]
s: 435
d: [2,524,1,7,3,0,40,130,140]
a: [3,529,1,2,1,524]
s: 529
d: [2,463,1,3,1,0,40,146,156]
a: [3,468,1,3,1,463]
s: 468
d: [2,495,2,2,1,1,40,332,136]
a: [3,501,2,2,1,495]
s: 501
d: [2,483,1,8,3,0,40,145,140]
a: [3,489,1,3,1,483]
s: 489
d: [2,548,1,5,2,1,40,149,306]
d: [2,558,1,3,1,2,21,548,220]
a: [3,562,1,3,1,558]
s: 562
d: [2,510,2,2,1,1,40,344,144]
a: [3,516,2,2,1,510]
s: 516
d: [2,472,1,5,2,2,44,306,449]
s: 472
d: [2,511,2,2,1,1,40,344,143]
a: [3,521,2,2,1,511]
s: 521
d: [2,482,1,5,2,0,40,145,156]
s: 482
d: [2,503,1,5,1,0,40,150,156]
s: 503
d: [2,534,2,2,1,1,40,349,144]
a: [3,539,2,2,1,534]
s: 539
d: [2,581,1,9,4,1,40,142,234]
d: [2,582,1,7,3,1,18,581]
d: [2,583,1,5,2,1,21,582,137]
s: 583
d: [2,535,2,2,1,1,40,349,143]
a: [3,544,2,2,1,535]
s: 544
d: [2,478,1,8,3,0,44,138,139]
d: [2,479,1,6,2,0,18,478]
s: 479
d: [2,759,2,4,2,1,44,663,479]
d: [2,768,2,2,1,2,21,759,388]
d: [2,769,2,0,0,3,21,768,195]
e: 769
r: 663
a: [3,662,1,3,1,658]
s: 662
d: [2,777,2,3,1,1,40,662,144]
a: [3,783,2,3,1,777]
s: 783
d: [2,477,1,8,3,0,44,139,138]
d: [2,480,1,6,2,0,18,477]
s: 480
d: [2,778,2,3,1,1,40,662,143]
a: [3,788,2,3,1,778]
s: 788
d: [2,504,1,6,2,0,40,150,140]
s: 504
d: [2,525,1,6,2,0,40,130,156]
s: 525
d: [2,758,2,6,2,1,44,184,479]
d: [2,762,2,4,1,2,21,758,195]
a: [3,766,2,4,1,762]
s: 766
d: [2,797,3,5,2,1,44,766,479]
d: [2,798,3,2,1,2,21,797,651]
d: [2,799,3,0,0,3,21,798,190]
e: 799
r: 179
r: 184
r: 607
r: 627
r: 766
r: 616
r: 651
s: 180
s: 608
s: 183
s: 615
a: [3,667,1,3,1,658]
s: 667
d: [2,828,2,4,2,1,44,667,479]
d: [2,829,2,2,1,2,21,828,393]
d: [2,830,2,0,0,3,21,829,190]
e: 830
r: 667
a: [3,666,1,3,1,658]
s: 666
s: 628
a: [3,671,1,5,1,658]
s: 671
d: [2,549,1,6,2,0,40,149,156]
s: 549
a: [3,650,2,3,1,647]
s: 650
d: [2,470,1,7,3,1,44,449,449]
s: 470
d: [2,832,2,3,1,1,40,666,144]
a: [3,838,2,3,1,832]
s: 838
d: [2,523,1,7,3,0,40,130,151]
s: 523
d: [2,547,1,10,3,0,40,149,140]
a: [3,552,1,7,2,547]
s: 552
d: [2,833,2,3,1,1,40,666,143]
a: [3,843,2,3,1,833]
s: 843
a: [3,555,1,3,1,547]
s: 555
d: [2,851,2,11,3,1,40,671,142]
a: [3,857,2,3,1,851]
s: 857
d: [2,907,3,4,2,1,44,857,479]
d: [2,908,3,2,1,1,18,907]
d: [2,909,3,0,0,2,21,908,195]
e: 909
r: 857
d: [2,566,1,7,2,0,40,155,156]
s: 566
d: [2,573,1,9,4,1,40,141,227]
d: [2,574,1,7,3,1,18,573]
s: 574
a: [3,856,2,3,1,851]
s: 856
d: [2,509,1,8,3,0,40,151,145]
s: 509
a: [3,861,2,3,1,851]
s: 861
d: [2,924,3,4,2,1,44,861,479]
d: [2,925,3,2,1,1,18,924]
d: [2,926,3,0,0,2,21,925,190]
e: 926
r: 861
d: [2,564,1,8,3,0,40,155,151]
s: 564
r: 552
r: 555
a: [3,556,1,3,1,547]
s: 556
d: [2,935,2,0,0,2,52,489,556,134]
e: 935
d: [2,936,2,3,1,1,40,556,134]
d: [2,940,2,0,0,2,21,936,489]
e: 940
r: 556
a: [3,860,2,3,1,851]
s: 860
s: 552
d: [2,886,2,6,2,1,40,523,183]
a: [3,895,2,3,1,886]
s: 895
d: [2,586,1,8,2,0,40,154,156]
s: 586
d: [2,584,1,9,3,0,40,154,151]
s: 584
d: [2,958,3,4,2,1,44,895,479]
d: [2,959,3,2,1,2,21,958,190]
a: [3,963,3,2,1,959]
s: 963
d: [2,585,1,9,3,0,40,154,140]
s: 585
d: [2,693,3,7,3,1,40,628,163]
d: [2,695,3,5,2,2,21,693,388]
d: [2,696,3,3,1,3,21,695,393]
a: [3,700,3,3,1,696]
s: 700
d: [2,546,1,10,3,0,40,149,151]
s: 546
d: [2,635,1,14,6,0,40,162,142]
d: [2,638,1,12,5,0,21,635,137]
d: [2,639,1,10,4,0,21,638,137]
s: 639
d: [2,972,4,4,2,1,40,700,523]
d: [2,974,4,2,1,2,21,972,963]
a: [3,978,4,2,1,974]
s: 978
d: [2,642,1,14,6,0,40,163,162]
d: [2,644,1,10,4,0,18,642]
s: 644
d: [2,917,3,3,1,1,40,856,136]
a: [3,922,3,3,1,917]
s: 922
d: [2,600,1,14,6,0,40,160,141]
s: 600
a: [3,765,2,4,1,762]
s: 765
d: [2,1006,3,5,2,1,44,765,479]
d: [2,1010,3,2,1,2,21,1006,650]
d: [2,1011,3,0,0,3,21,1010,190]
e: 1011
i: [1,123,0,8,2,1,0,0] t4_algseq_1
d: [2,702,1,7,3,1,40,179,141]
d: [2,703,1,5,2,2,21,702,190]
d: [2,704,1,3,1,3,21,703,195]
d: [2,707,2,7,3,1,40,616,163]
d: [2,709,2,5,2,2,21,707,393]
d: [2,710,2,3,1,3,21,709,388]
d: [2,708,2,7,3,1,40,616,142]
d: [2,801,2,5,2,2,21,708,190]
d: [2,802,2,3,1,3,21,801,195]
d: [2,713,3,7,3,1,40,651,164]
d: [2,804,3,5,2,2,21,713,393]
d: [2,805,3,3,1,3,21,804,388]
d: [2,816,2,6,2,1,44,183,479]
d: [2,820,2,4,1,2,21,816,195]
f: 1011,820,805,802,799,762,710,704,382,381,121,122,123,124
% Refutation found. Thanks to Tanya!
% SZS status Theorem for algseq_1__t4_algseq_1
% SZS output start Proof for algseq_1__t4_algseq_1
1. ! [X0] : (v7_ordinal1(X0) => ! [X1] : (v7_ordinal1(X1) => (r1_xxreal_0(X0,X1) <=> r1_tarski(k7_card_1(X0),k7_card_1(X1))))) [input t4_algseq_1]
2. ~! [X0] : (v7_ordinal1(X0) => ! [X1] : (v7_ordinal1(X1) => (r1_xxreal_0(X0,X1) <=> r1_tarski(k7_card_1(X0),k7_card_1(X1))))) [negated conjecture 1]
37. ! [X0] : (v7_ordinal1(X0) => k6_card_1(X0) = X0) [input d6_card_1]
43. ! [X0] : (v7_ordinal1(X0) => k7_card_1(X0) = k6_card_1(X0)) [input redefinition_k7_card_1]
44. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
47. ! [X0] : (v7_ordinal1(X0) => v3_ordinal1(X0)) [input cc6_ordinal1]
50. ! [X0] : (v7_ordinal1(X0) => ! [X1] : (v7_ordinal1(X1) => (r1_xxreal_0(X0,X1) <=> r1_ordinal1(X0,X1)))) [input t39_nat_1]
62. ? [X0] : (? [X1] : ((r1_xxreal_0(X0,X1) <~> r1_tarski(k7_card_1(X0),k7_card_1(X1))) & v7_ordinal1(X1)) & v7_ordinal1(X0)) [ennf transformation 2]
70. ! [X0] : (v3_ordinal1(X0) | ~v7_ordinal1(X0)) [ennf transformation 47]
71. ! [X0] : (k6_card_1(X0) = X0 | ~v7_ordinal1(X0)) [ennf transformation 37]
72. ! [X0] : (k7_card_1(X0) = k6_card_1(X0) | ~v7_ordinal1(X0)) [ennf transformation 43]
74. ! [X0] : (! [X1] : ((r1_xxreal_0(X0,X1) <=> r1_ordinal1(X0,X1)) | ~v7_ordinal1(X1)) | ~v7_ordinal1(X0)) [ennf transformation 50]
97. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 44]
98. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 97]
99. ? [X0] : (? [X1] : (((~r1_tarski(k7_card_1(X0),k7_card_1(X1)) | ~r1_xxreal_0(X0,X1)) & (r1_tarski(k7_card_1(X0),k7_card_1(X1)) | r1_xxreal_0(X0,X1))) & v7_ordinal1(X1)) & v7_ordinal1(X0)) [nnf transformation 62]
100. ? [X0] : (? [X1] : ((~r1_tarski(k7_card_1(X0),k7_card_1(X1)) | ~r1_xxreal_0(X0,X1)) & (r1_tarski(k7_card_1(X0),k7_card_1(X1)) | r1_xxreal_0(X0,X1)) & v7_ordinal1(X1)) & v7_ordinal1(X0)) [flattening 99]
101. ? [X0] : (? [X1] : ((~r1_tarski(k7_card_1(X0),k7_card_1(X1)) | ~r1_xxreal_0(X0,X1)) & (r1_tarski(k7_card_1(X0),k7_card_1(X1)) | r1_xxreal_0(X0,X1)) & v7_ordinal1(X1)) & v7_ordinal1(X0)) => (? [X1] : ((~r1_tarski(k7_card_1(sK0),k7_card_1(X1)) | ~r1_xxreal_0(sK0,X1)) & (r1_tarski(k7_card_1(sK0),k7_card_1(X1)) | r1_xxreal_0(sK0,X1)) & v7_ordinal1(X1)) & v7_ordinal1(sK0)) [choice axiom]
102. ? [X1] : ((~r1_tarski(k7_card_1(sK0),k7_card_1(X1)) | ~r1_xxreal_0(sK0,X1)) & (r1_tarski(k7_card_1(sK0),k7_card_1(X1)) | r1_xxreal_0(sK0,X1)) & v7_ordinal1(X1)) => ((~r1_tarski(k7_card_1(sK0),k7_card_1(sK1)) | ~r1_xxreal_0(sK0,sK1)) & (r1_tarski(k7_card_1(sK0),k7_card_1(sK1)) | r1_xxreal_0(sK0,sK1)) & v7_ordinal1(sK1)) [choice axiom]
103. ((~r1_tarski(k7_card_1(sK0),k7_card_1(sK1)) | ~r1_xxreal_0(sK0,sK1)) & (r1_tarski(k7_card_1(sK0),k7_card_1(sK1)) | r1_xxreal_0(sK0,sK1)) & v7_ordinal1(sK1)) & v7_ordinal1(sK0) [skolemisation 100,102,101]
104. ! [X0] : (! [X1] : (((r1_xxreal_0(X0,X1) | ~r1_ordinal1(X0,X1)) & (r1_ordinal1(X0,X1) | ~r1_xxreal_0(X0,X1))) | ~v7_ordinal1(X1)) | ~v7_ordinal1(X0)) [nnf transformation 74]
108. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 98]
121. v7_ordinal1(sK0) [cnf transformation 103] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
122. v7_ordinal1(sK1) [cnf transformation 103] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
123. r1_tarski(k7_card_1(sK0),k7_card_1(sK1)) | r1_xxreal_0(sK0,sK1) [cnf transformation 103] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
124. ~r1_tarski(k7_card_1(sK0),k7_card_1(sK1)) | ~r1_xxreal_0(sK0,sK1) [cnf transformation 103] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
137. ~v7_ordinal1(X0) | v3_ordinal1(X0) [cnf transformation 70] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
138. k6_card_1(X0) = X0 | ~v7_ordinal1(X0) [cnf transformation 71] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
139. k7_card_1(X0) = k6_card_1(X0) | ~v7_ordinal1(X0) [cnf transformation 72] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
141. ~r1_xxreal_0(X0,X1) | r1_ordinal1(X0,X1) | ~v7_ordinal1(X1) | ~v7_ordinal1(X0) [cnf transformation 104] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
142. ~r1_ordinal1(X0,X1) | r1_xxreal_0(X0,X1) | ~v7_ordinal1(X1) | ~v7_ordinal1(X0) [cnf transformation 104] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
163. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 108] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
164. ~r1_tarski(X0,X1) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 108] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
178. 1 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
179. r1_xxreal_0(sK0,sK1) <- (1) [avatar component clause 178] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
182. 2 <=> r1_tarski(k7_card_1(sK0),k7_card_1(sK1)) [avatar definition]
183. r1_tarski(k7_card_1(sK0),k7_card_1(sK1)) <- (2) [avatar component clause 182] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
184. ~r1_tarski(k7_card_1(sK0),k7_card_1(sK1)) <- (~2) [avatar component clause 182] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
185. ~1 | ~2 [avatar split clause 124,182,178]
186. 1 | 2 [avatar split clause 123,182,178]
188. 3 <=> v7_ordinal1(sK1) [avatar definition]
190. v7_ordinal1(sK1) <- (3) [avatar component clause 188] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
191. 3 [avatar split clause 122,188]
193. 4 <=> v7_ordinal1(sK0) [avatar definition]
195. v7_ordinal1(sK0) <- (4) [avatar component clause 193] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
196. 4 [avatar split clause 121,193]
381. v3_ordinal1(sK0) <- (4) [resolution 137,195] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
382. v3_ordinal1(sK1) <- (3) [resolution 137,190] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
386. 39 <=> v3_ordinal1(sK0) [avatar definition]
388. v3_ordinal1(sK0) <- (39) [avatar component clause 386] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
389. 39 | ~4 [avatar split clause 381,193,386]
391. 40 <=> v3_ordinal1(sK1) [avatar definition]
393. v3_ordinal1(sK1) <- (40) [avatar component clause 391] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
394. 40 | ~3 [avatar split clause 382,188,391]
478. k7_card_1(X0) = X0 | ~v7_ordinal1(X0) | ~v7_ordinal1(X0) [superposition 138,139] {a:1,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
479. k7_card_1(X0) = X0 | ~v7_ordinal1(X0) [duplicate literal removal 478] {a:1,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
614. 70 <=> r1_ordinal1(sK0,sK1) [avatar definition]
616. r1_ordinal1(sK0,sK1) <- (70) [avatar component clause 614] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
649. 72 <=> r1_tarski(sK0,sK1) [avatar definition]
650. ~r1_tarski(sK0,sK1) <- (~72) [avatar component clause 649] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
651. r1_tarski(sK0,sK1) <- (72) [avatar component clause 649] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
702. r1_ordinal1(sK0,sK1) | ~v7_ordinal1(sK1) | ~v7_ordinal1(sK0) <- (1) [resolution 179,141] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
703. r1_ordinal1(sK0,sK1) | ~v7_ordinal1(sK0) <- (1, 3) [subsumption resolution 702,190] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
704. r1_ordinal1(sK0,sK1) <- (1, 3, 4) [subsumption resolution 703,195] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
705. 70 | ~1 | ~3 | ~4 [avatar split clause 704,193,188,178,614]
707. r1_tarski(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (70) [resolution 616,163] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
708. r1_xxreal_0(sK0,sK1) | ~v7_ordinal1(sK1) | ~v7_ordinal1(sK0) <- (70) [resolution 616,142] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
709. r1_tarski(sK0,sK1) | ~v3_ordinal1(sK0) <- (40, 70) [subsumption resolution 707,393] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
710. r1_tarski(sK0,sK1) <- (39, 40, 70) [subsumption resolution 709,388] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
711. 72 | ~39 | ~40 | ~70 [avatar split clause 710,614,391,386,649]
713. r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (72) [resolution 651,164] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
758. ~r1_tarski(sK0,k7_card_1(sK1)) | ~v7_ordinal1(sK0) <- (~2) [superposition 184,479] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
762. ~r1_tarski(sK0,k7_card_1(sK1)) <- (~2, 4) [subsumption resolution 758,195] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
764. 83 <=> r1_tarski(sK0,k7_card_1(sK1)) [avatar definition]
765. r1_tarski(sK0,k7_card_1(sK1)) <- (83) [avatar component clause 764] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
766. ~r1_tarski(sK0,k7_card_1(sK1)) <- (~83) [avatar component clause 764] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
767. ~83 | 2 | ~4 [avatar split clause 762,193,182,764]
797. ~r1_tarski(sK0,sK1) | ~v7_ordinal1(sK1) <- (~83) [superposition 766,479] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
798. ~v7_ordinal1(sK1) <- (72, ~83) [subsumption resolution 797,651] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
799. $false <- (3, 72, ~83) [subsumption resolution 798,190] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
800. ~3 | ~72 | 83 [avatar contradiction clause 799]
801. r1_xxreal_0(sK0,sK1) | ~v7_ordinal1(sK0) <- (3, 70) [subsumption resolution 708,190] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
802. r1_xxreal_0(sK0,sK1) <- (3, 4, 70) [subsumption resolution 801,195] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
803. 1 | ~3 | ~4 | ~70 [avatar split clause 802,614,193,188,178]
804. r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK0) <- (40, 72) [subsumption resolution 713,393] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
805. r1_ordinal1(sK0,sK1) <- (39, 40, 72) [subsumption resolution 804,388] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
806. 70 | ~39 | ~40 | ~72 [avatar split clause 805,649,391,386,614]
816. r1_tarski(sK0,k7_card_1(sK1)) | ~v7_ordinal1(sK0) <- (2) [superposition 183,479] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
820. r1_tarski(sK0,k7_card_1(sK1)) <- (2, 4) [subsumption resolution 816,195] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
821. 83 | ~2 | ~4 [avatar split clause 820,193,182,764]
1006. r1_tarski(sK0,sK1) | ~v7_ordinal1(sK1) <- (83) [superposition 765,479] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1010. ~v7_ordinal1(sK1) <- (~72, 83) [subsumption resolution 1006,650] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1011. $false <- (3, ~72, 83) [subsumption resolution 1010,190] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1012. ~3 | 72 | ~83 [avatar contradiction clause 1011]
1013. $false [avatar sat refutation 185,186,191,196,389,394,705,711,767,800,803,806,821,1012] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:742,thDist:-742}
% SZS output end Proof for algseq_1__t4_algseq_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 49
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 56
% Generated clauses: 499
% Active clauses: 160
% Passive clauses: 218
% Final active clauses: 129
% Final passive clauses: 48
% Main loop iterations started: 166
% 
% >>> Simplifying Inferences
% Duplicate literals: 27
% Fw subsumption resolutions: 70
% Fw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 58
% 
% >>> Generating Inferences
% Binary resolution: 175
% Unit resulting resolution: 2
% Factoring: 2
% Forward superposition: 30
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 15
% Split components: 32
% Sat splitting refutations: 11
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 128
% SAT solver unit clauses: 20
% SAT solver binary clauses: 64
% 
% Memory used [KB]: 5628
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% SAT solver time: 0.001 s
% congruence closure: 0.001 s
% superposition: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

