i: [1,93,0,2,1,1,0,0] t2_tops_3
a: [3,170,0,2,1,93]
s: 170
i: [1,110,0,21,4,1,0,0] t2_tops_3
a: [3,148,0,4,1,110]
s: 148
i: [1,94,0,5,1,1,0,0] t2_tops_3
a: [3,165,0,5,1,94]
s: 165
a: [3,145,0,7,1,110]
s: 145
i: [1,111,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,175,0,2,1,111]
s: 175
i: [1,140,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,185,0,2,1,140]
s: 185
i: [1,141,0,2,1,0,0,2] existence_l1_struct_0
a: [3,190,0,2,1,141]
s: 190
i: [1,136,0,3,1,0,0,255] reflexivity_r1_tarski
s: 136
i: [1,112,0,3,1,0,0,255] d2_xboole_0
a: [3,180,0,3,1,112]
s: 180
r: 175
d: [2,192,0,2,1,2,23,175,180]
a: [3,196,0,2,1,192]
s: 196
i: [1,115,0,4,2,0,0,1] dt_l1_pre_topc
s: 115
d: [2,198,1,2,1,1,40,115,170]
a: [3,203,1,2,1,198]
s: 203
i: [1,135,0,4,1,0,0,1] existence_m1_subset_1
s: 135
d: [2,199,1,2,1,1,40,115,185]
a: [3,208,1,2,1,199]
s: 208
i: [1,113,0,5,1,0,0,2] t4_boole
s: 113
i: [1,114,0,5,1,0,0,2] t3_boole
s: 114
i: [1,118,0,5,2,0,0,1] fc3_struct_0
s: 118
i: [1,123,0,5,2,0,0,1] t6_boole
s: 123
i: [1,132,0,5,2,0,0,255] t7_boole
s: 132
i: [1,119,0,6,2,0,0,1] d2_struct_0
s: 119
i: [1,124,0,6,2,0,0,255] t1_subset
s: 124
i: [1,137,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 137
i: [1,120,0,7,2,0,0,1] d3_struct_0
s: 120
r: 145
r: 148
a: [3,149,0,4,1,110]
s: 149
a: [3,157,0,4,1,110]
s: 157
a: [3,144,0,7,1,110]
s: 144
i: [1,127,0,7,2,0,0,255] t3_subset
s: 127
i: [1,128,0,7,2,0,0,255] t3_subset
s: 128
d: [2,225,1,4,1,1,40,127,165]
a: [3,230,1,4,1,225]
s: 230
i: [1,131,0,7,3,0,0,2] t8_boole
s: 131
i: [1,116,0,8,2,0,0,255] rc4_tops_1
s: 116
d: [2,226,1,5,1,0,40,127,135]
s: 226
i: [1,117,0,8,2,0,0,255] rc5_tops_1
s: 117
d: [2,210,1,7,3,0,44,123,123]
s: 210
i: [1,121,0,8,2,0,0,1] dt_k1_struct_0
s: 121
i: [1,122,0,8,2,0,0,1] dt_k2_struct_0
s: 122
d: [2,211,1,7,2,0,44,114,123]
s: 211
i: [1,125,0,8,3,0,0,255] t2_subset
s: 125
d: [2,244,1,9,2,1,40,125,165]
a: [3,251,1,5,1,244]
s: 251
i: [1,129,0,8,2,0,0,255] t37_xboole_1
s: 129
i: [1,130,0,8,2,0,0,255] t37_xboole_1
s: 130
a: [3,254,1,4,1,244]
s: 254
i: [1,134,0,9,3,0,0,255] t5_subset
s: 134
d: [2,278,1,6,2,1,40,134,165]
a: [3,288,1,3,1,278]
s: 288
i: [1,133,0,10,3,0,0,255] t4_subset
s: 133
i: [1,138,0,10,2,0,0,1] dt_k3_subset_1
s: 138
r: 144
r: 149
r: 157
s: 148
a: [3,156,0,4,1,110]
s: 156
r: 165
r: 230
r: 251
r: 254
r: 288
d: [2,319,1,2,1,2,23,288,156]
a: [3,342,1,2,1,319]
s: 342
d: [2,315,0,4,1,2,23,165,156]
a: [3,327,0,4,1,315]
s: 327
s: 145
d: [2,318,1,3,1,2,23,254,156]
a: [3,336,1,3,1,318]
s: 336
d: [2,374,0,6,1,2,22,145,156]
a: [3,378,0,6,1,374]
s: 378
d: [2,317,1,4,1,2,23,251,156]
a: [3,331,1,4,1,317]
s: 331
i: [1,126,0,11,2,0,0,1] d4_subset_1
s: 126
i: [1,139,0,11,2,0,0,1] involutiveness_k3_subset_1
s: 139
d: [2,321,1,7,2,1,44,121,156]
d: [2,347,1,5,1,2,21,321,203]
a: [3,351,1,5,1,347]
s: 351
d: [2,322,1,7,2,1,44,117,156]
d: [2,353,1,5,1,2,21,322,170]
a: [3,357,1,5,1,353]
s: 357
d: [2,392,2,4,1,1,40,351,127]
a: [3,398,2,4,1,392]
s: 398
d: [2,323,1,7,2,1,44,116,156]
d: [2,359,1,5,1,2,21,323,170]
a: [3,363,1,5,1,359]
s: 363
d: [2,293,1,7,2,1,40,133,165]
s: 293
d: [2,415,3,5,2,1,44,398,119]
d: [2,416,3,3,1,2,21,415,203]
a: [3,420,3,3,1,416]
s: 420
d: [2,266,1,6,2,0,44,129,113]
d: [2,269,1,3,1,0,19,266]
s: 269
r: 420
d: [2,394,2,6,2,1,44,351,119]
d: [2,400,2,4,1,2,21,394,203]
a: [3,404,2,4,1,400]
s: 404
d: [2,247,1,6,2,0,40,125,135]
s: 247
d: [2,267,1,6,2,0,44,129,114]
s: 267
d: [2,408,2,4,1,1,40,357,127]
a: [3,413,2,4,1,408]
s: 413
d: [2,270,1,6,2,0,44,130,114]
s: 270
d: [2,424,2,4,1,1,40,363,127]
a: [3,429,2,4,1,424]
s: 429
d: [2,272,1,6,2,0,44,114,130]
s: 272
d: [2,212,1,7,2,0,44,113,123]
s: 212
d: [2,257,2,5,1,1,40,251,137]
a: [3,263,2,5,1,257]
s: 263
d: [2,233,1,7,2,0,40,116,127]
s: 233
d: [2,459,2,4,1,2,22,263,156]
a: [3,463,2,4,1,459]
s: 463
d: [2,234,1,7,2,0,40,117,127]
s: 234
d: [2,235,1,7,2,0,40,121,127]
s: 235
d: [2,432,4,5,2,1,44,420,123]
s: 432
d: [2,236,1,9,3,0,44,121,119]
d: [2,237,1,7,2,0,18,236]
s: 237
d: [2,434,3,6,2,1,40,404,133]
s: 434
d: [2,238,1,7,2,0,40,122,127]
s: 238
d: [2,281,1,7,2,0,40,134,135]
s: 281
d: [2,438,3,6,2,1,44,404,123]
s: 438
d: [2,239,1,10,3,0,44,122,120]
d: [2,240,1,8,2,0,18,239]
s: 240
d: [2,390,2,7,2,1,40,351,133]
s: 390
d: [2,275,1,8,3,0,40,134,128]
s: 275
d: [2,296,1,8,2,0,40,133,135]
s: 296
d: [2,503,3,8,2,1,40,390,247]
a: [3,508,3,3,1,503]
s: 508
d: [2,276,1,9,3,0,40,134,122]
s: 276
d: [2,406,2,7,2,1,40,357,133]
s: 406
d: [2,277,1,9,3,0,40,134,121]
s: 277
d: [2,279,1,9,3,0,40,134,116]
s: 279
d: [2,537,3,8,2,1,40,406,247]
a: [3,541,3,3,1,537]
s: 541
d: [2,280,1,9,3,0,40,134,117]
s: 280
d: [2,422,2,7,2,1,40,363,133]
s: 422
d: [2,290,1,9,3,0,40,133,128]
s: 290
d: [2,299,1,9,2,0,40,138,127]
s: 299
d: [2,553,3,8,2,1,40,422,247]
a: [3,557,3,3,1,553]
s: 557
d: [2,380,1,9,2,0,44,126,113]
s: 380
d: [2,471,2,8,3,0,40,237,134]
a: [3,478,2,3,1,471]
s: 478
r: 434
d: [2,381,1,9,2,0,44,126,114]
s: 381
d: [2,384,1,9,2,0,44,113,126]
s: 384
d: [2,582,2,12,3,0,44,138,380]
d: [2,584,2,8,2,0,18,582]
a: [3,590,2,4,1,584]
s: 590
r: 380
d: [2,611,3,0,0,1,40,590,135]
e: 611
d: [2,606,3,0,0,1,52,135,590]
e: 606
d: [2,607,3,0,0,1,52,136,590,128]
e: 607
d: [2,608,3,0,0,1,52,135,590,138]
e: 608
r: 590
s: 380
a: [3,594,2,4,1,584]
s: 594
d: [2,385,1,9,2,0,44,114,126]
s: 385
d: [2,241,1,10,3,0,40,125,128]
s: 241
d: [2,632,3,7,2,1,40,594,125]
a: [3,637,3,4,1,632]
s: 637
d: [2,291,1,10,3,0,40,133,122]
s: 291
a: [3,640,3,3,1,632]
s: 640
d: [2,292,1,10,3,0,40,133,121]
s: 292
d: [2,294,1,10,3,0,40,133,116]
s: 294
d: [2,433,2,5,2,0,44,269,123]
s: 433
r: 432
d: [2,295,1,10,3,0,40,133,117]
s: 295
d: [2,446,2,5,1,0,40,270,226]
a: [3,451,2,5,1,446]
s: 451
d: [2,298,1,11,3,0,40,138,134]
s: 298
d: [2,242,1,12,3,0,40,125,122]
s: 242
d: [2,457,2,8,3,0,44,129,212]
d: [2,458,2,5,2,0,21,457,123]
s: 458
d: [2,243,1,12,3,0,40,125,121]
s: 243
d: [2,658,4,5,2,1,44,640,123]
s: 658
d: [2,245,1,12,3,0,40,125,116]
s: 245
d: [2,246,1,12,3,0,40,125,117]
s: 246
d: [2,439,2,6,2,0,40,247,137]
s: 439
d: [2,297,1,12,3,0,40,138,133]
s: 297
d: [2,493,2,6,2,0,40,281,247]
s: 493
d: [2,382,1,12,3,0,44,126,130]
s: 382
d: [2,386,1,12,3,0,44,130,126]
s: 386
d: [2,722,2,11,3,1,44,378,382]
d: [2,732,2,7,2,2,21,722,327]
d: [2,733,2,4,1,2,21,732,136]
a: [3,737,2,4,1,733]
s: 737
d: [2,739,4,0,0,2,52,508,737,123]
e: 739
d: [2,740,4,0,0,3,52,196,508,737,131]
e: 740
d: [2,741,3,0,0,2,52,203,737,119]
e: 741
d: [2,742,4,0,0,3,52,196,508,737,131]
e: 742
d: [2,743,3,5,2,1,44,737,119]
d: [2,744,3,2,1,1,19,743]
d: [2,749,3,0,0,2,21,744,203]
e: 749
r: 148
r: 156
r: 327
r: 737
s: 149
s: 288
s: 157
a: [3,326,0,4,1,315]
s: 326
d: [2,757,1,0,0,1,52,136,326,128]
e: 757
d: [2,759,1,3,1,1,40,326,128]
d: [2,761,1,0,0,1,21,759,136]
e: 761
r: 326
r: 378
s: 327
s: 165
s: 230
a: [3,377,0,6,1,374]
s: 377
s: 144
a: [3,255,1,4,1,244]
s: 255
a: [3,250,1,5,1,244]
s: 250
a: [3,736,2,4,1,733]
s: 736
r: 144
r: 377
r: 508
r: 390
r: 398
r: 351
d: [2,827,0,5,1,2,23,377,736]
a: [3,842,0,5,1,827]
s: 842
a: [3,152,0,6,1,110]
s: 152
d: [2,847,1,9,2,1,44,139,842]
d: [2,849,1,5,1,2,21,847,327]
a: [3,853,1,5,1,849]
s: 853
d: [2,859,1,10,2,1,44,138,152]
d: [2,866,1,5,1,2,21,859,165]
a: [3,870,1,5,1,866]
s: 870
s: 263
d: [2,301,1,11,2,1,44,138,144]
d: [2,302,1,6,1,2,21,301,165]
a: [3,306,1,6,1,302]
s: 306
d: [2,780,1,10,2,1,44,139,377]
d: [2,788,1,6,1,2,21,780,327]
a: [3,792,1,6,1,788]
s: 792
d: [2,798,2,10,2,1,44,299,144]
d: [2,810,2,5,1,2,21,798,165]
a: [3,814,2,5,1,810]
s: 814
d: [2,858,1,11,2,1,44,139,152]
d: [2,860,1,6,1,2,21,858,165]
a: [3,864,1,6,1,860]
s: 864
d: [2,889,2,9,2,1,44,864,381]
d: [2,898,2,5,1,2,21,889,157]
d: [2,899,2,0,0,3,21,898,870]
e: 899
d: [2,891,2,9,2,1,44,381,864]
d: [2,902,2,5,1,2,21,891,157]
d: [2,903,2,0,0,3,21,902,870]
e: 903
i: [1,95,0,21,4,1,0,0] t2_tops_3
d: [2,310,1,6,2,1,44,120,148]
d: [2,313,1,4,1,2,21,310,203]
d: [2,752,2,8,2,1,21,722,136]
d: [2,754,1,6,2,1,44,149,120]
d: [2,755,1,4,1,2,21,754,203]
d: [2,826,0,6,1,2,23,144,736]
f: 903,866,860,826,761,755,752,749,374,313,198,93,94,95,110
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tops_3__t2_tops_3
% SZS output start Proof for tops_3__t2_tops_3
1. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ((k3_subset_1(u1_struct_0(X0),X1) = k1_xboole_0 => u1_struct_0(X0) = X1) & (u1_struct_0(X0) = X1 => k3_subset_1(u1_struct_0(X0),X1) = k1_xboole_0) & (k3_subset_1(u1_struct_0(X0),X1) = k1_struct_0(X0) => k2_struct_0(X0) = X1) & (k2_struct_0(X0) = X1 => k3_subset_1(u1_struct_0(X0),X1) = k1_struct_0(X0))))) [input t2_tops_3]
2. ~! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ((k3_subset_1(u1_struct_0(X0),X1) = k1_xboole_0 => u1_struct_0(X0) = X1) & (u1_struct_0(X0) = X1 => k3_subset_1(u1_struct_0(X0),X1) = k1_xboole_0) & (k3_subset_1(u1_struct_0(X0),X1) = k1_struct_0(X0) => k2_struct_0(X0) = X1) & (k2_struct_0(X0) = X1 => k3_subset_1(u1_struct_0(X0),X1) = k1_struct_0(X0))))) [negated conjecture 1]
5. ! [X0] : (l1_struct_0(X0) => k1_struct_0(X0) = k1_xboole_0) [input d2_struct_0]
7. ! [X0] : (l1_struct_0(X0) => u1_struct_0(X0) = k2_struct_0(X0)) [input d3_struct_0]
8. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => k3_subset_1(X0,X1) = k4_xboole_0(X0,X1)) [input d4_subset_1]
13. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => m1_subset_1(k3_subset_1(X0,X1),k1_zfmisc_1(X0))) [input dt_k3_subset_1]
15. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
26. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => k3_subset_1(X0,k3_subset_1(X0,X1)) = X1) [input involutiveness_k3_subset_1]
29. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
32. ! [X0] : ! [X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [input t37_xboole_1]
33. ! [X0] : k4_xboole_0(X0,k1_xboole_0) = X0 [input t3_boole]
34. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
43. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => k3_subset_1(X0,X1) = k4_xboole_0(X0,X1)) [flattening 8]
44. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 34]
45. ! [X0,X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [flattening 32]
50. ! [X0] : r1_tarski(X0,X0) [rectify 29]
56. ? [X0] : (? [X1] : (((u1_struct_0(X0) != X1 & k3_subset_1(u1_struct_0(X0),X1) = k1_xboole_0) | (k3_subset_1(u1_struct_0(X0),X1) != k1_xboole_0 & u1_struct_0(X0) = X1) | (k2_struct_0(X0) != X1 & k3_subset_1(u1_struct_0(X0),X1) = k1_struct_0(X0)) | (k3_subset_1(u1_struct_0(X0),X1) != k1_struct_0(X0) & k2_struct_0(X0) = X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0)) [ennf transformation 2]
57. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 15]
61. ! [X0] : (k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0)) [ennf transformation 5]
62. ! [X0] : (u1_struct_0(X0) = k2_struct_0(X0) | ~l1_struct_0(X0)) [ennf transformation 7]
69. ! [X0,X1] : (k3_subset_1(X0,X1) = k4_xboole_0(X0,X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 43]
76. ! [X0,X1] : (m1_subset_1(k3_subset_1(X0,X1),k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 13]
77. ! [X0,X1] : (k3_subset_1(X0,k3_subset_1(X0,X1)) = X1 | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 26]
78. ? [X0] : (? [X1] : (((u1_struct_0(X0) != X1 & k3_subset_1(u1_struct_0(X0),X1) = k1_xboole_0) | (k3_subset_1(u1_struct_0(X0),X1) != k1_xboole_0 & u1_struct_0(X0) = X1) | (k2_struct_0(X0) != X1 & k3_subset_1(u1_struct_0(X0),X1) = k1_struct_0(X0)) | (k3_subset_1(u1_struct_0(X0),X1) != k1_struct_0(X0) & k2_struct_0(X0) = X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0)) => (? [X1] : (((u1_struct_0(sK0) != X1 & k1_xboole_0 = k3_subset_1(u1_struct_0(sK0),X1)) | (k1_xboole_0 != k3_subset_1(u1_struct_0(sK0),X1) & u1_struct_0(sK0) = X1) | (k2_struct_0(sK0) != X1 & k3_subset_1(u1_struct_0(sK0),X1) = k1_struct_0(sK0)) | (k3_subset_1(u1_struct_0(sK0),X1) != k1_struct_0(sK0) & k2_struct_0(sK0) = X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0)) [choice axiom]
79. ? [X1] : (((u1_struct_0(sK0) != X1 & k1_xboole_0 = k3_subset_1(u1_struct_0(sK0),X1)) | (k1_xboole_0 != k3_subset_1(u1_struct_0(sK0),X1) & u1_struct_0(sK0) = X1) | (k2_struct_0(sK0) != X1 & k3_subset_1(u1_struct_0(sK0),X1) = k1_struct_0(sK0)) | (k3_subset_1(u1_struct_0(sK0),X1) != k1_struct_0(sK0) & k2_struct_0(sK0) = X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => (((u1_struct_0(sK0) != sK1 & k1_xboole_0 = k3_subset_1(u1_struct_0(sK0),sK1)) | (k1_xboole_0 != k3_subset_1(u1_struct_0(sK0),sK1) & u1_struct_0(sK0) = sK1) | (k2_struct_0(sK0) != sK1 & k1_struct_0(sK0) = k3_subset_1(u1_struct_0(sK0),sK1)) | (k1_struct_0(sK0) != k3_subset_1(u1_struct_0(sK0),sK1) & k2_struct_0(sK0) = sK1)) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
80. (((u1_struct_0(sK0) != sK1 & k1_xboole_0 = k3_subset_1(u1_struct_0(sK0),sK1)) | (k1_xboole_0 != k3_subset_1(u1_struct_0(sK0),sK1) & u1_struct_0(sK0) = sK1) | (k2_struct_0(sK0) != sK1 & k1_struct_0(sK0) = k3_subset_1(u1_struct_0(sK0),sK1)) | (k1_struct_0(sK0) != k3_subset_1(u1_struct_0(sK0),sK1) & k2_struct_0(sK0) = sK1)) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) [skolemisation 56,79,78]
85. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 44]
86. ! [X0,X1] : ((k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | k1_xboole_0 != k4_xboole_0(X0,X1))) [nnf transformation 45]
93. l1_pre_topc(sK0) [cnf transformation 80] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
94. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 80] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
95. k1_xboole_0 = k3_subset_1(u1_struct_0(sK0),sK1) | u1_struct_0(sK0) = sK1 | k1_struct_0(sK0) = k3_subset_1(u1_struct_0(sK0),sK1) | k2_struct_0(sK0) = sK1 [cnf transformation 80] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
110. u1_struct_0(sK0) != sK1 | k1_xboole_0 != k3_subset_1(u1_struct_0(sK0),sK1) | k2_struct_0(sK0) != sK1 | k1_struct_0(sK0) != k3_subset_1(u1_struct_0(sK0),sK1) [cnf transformation 80] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
114. k4_xboole_0(X0,k1_xboole_0) = X0 [cnf transformation 33] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
115. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 57] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
119. k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0) [cnf transformation 61] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. u1_struct_0(X0) = k2_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 62] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
126. k3_subset_1(X0,X1) = k4_xboole_0(X0,X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 69] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
128. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 85] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
130. k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1) [cnf transformation 86] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
136. r1_tarski(X0,X0) [cnf transformation 50] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
138. m1_subset_1(k3_subset_1(X0,X1),k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 76] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
139. k3_subset_1(X0,k3_subset_1(X0,X1)) = X1 | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 77] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
143. 1 <=> k1_struct_0(sK0) = k3_subset_1(u1_struct_0(sK0),sK1) [avatar definition]
144. k1_struct_0(sK0) = k3_subset_1(u1_struct_0(sK0),sK1) <- (1) [avatar component clause 143] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
145. k1_struct_0(sK0) != k3_subset_1(u1_struct_0(sK0),sK1) <- (~1) [avatar component clause 143] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
147. 2 <=> k2_struct_0(sK0) = sK1 [avatar definition]
148. k2_struct_0(sK0) = sK1 <- (2) [avatar component clause 147] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
149. k2_struct_0(sK0) != sK1 <- (~2) [avatar component clause 147] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
151. 3 <=> k1_xboole_0 = k3_subset_1(u1_struct_0(sK0),sK1) [avatar definition]
152. k1_xboole_0 = k3_subset_1(u1_struct_0(sK0),sK1) <- (3) [avatar component clause 151] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
155. 4 <=> u1_struct_0(sK0) = sK1 [avatar definition]
156. u1_struct_0(sK0) = sK1 <- (4) [avatar component clause 155] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
157. u1_struct_0(sK0) != sK1 <- (~4) [avatar component clause 155] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
158. ~1 | ~2 | ~3 | ~4 [avatar split clause 110,155,151,147,143]
161. 2 | 1 | 4 | 3 [avatar split clause 95,151,155,143,147]
163. 5 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
165. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (5) [avatar component clause 163] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
166. 5 [avatar split clause 94,163]
168. 6 <=> l1_pre_topc(sK0) [avatar definition]
170. l1_pre_topc(sK0) <- (6) [avatar component clause 168] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
171. 6 [avatar split clause 93,168]
198. l1_struct_0(sK0) <- (6) [resolution 115,170] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
201. 12 <=> l1_struct_0(sK0) [avatar definition]
203. l1_struct_0(sK0) <- (12) [avatar component clause 201] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
204. 12 | ~6 [avatar split clause 198,168,201]
310. u1_struct_0(sK0) = sK1 | ~l1_struct_0(sK0) <- (2) [superposition 120,148] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
313. u1_struct_0(sK0) = sK1 <- (2, 12) [subsumption resolution 310,203] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
314. 4 | ~2 | ~12 [avatar split clause 313,201,147,155]
325. 21 <=> m1_subset_1(sK1,k1_zfmisc_1(sK1)) [avatar definition]
326. ~m1_subset_1(sK1,k1_zfmisc_1(sK1)) <- (~21) [avatar component clause 325] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
374. k1_struct_0(sK0) != k3_subset_1(sK1,sK1) <- (~1, 4) [forward demodulation 145,156] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
376. 28 <=> k1_struct_0(sK0) = k3_subset_1(sK1,sK1) [avatar definition]
378. k1_struct_0(sK0) != k3_subset_1(sK1,sK1) <- (~28) [avatar component clause 376] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
379. ~28 | 1 | ~4 [avatar split clause 374,155,143,376]
381. k3_subset_1(X1,k1_xboole_0) = X1 | ~m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)) [superposition 126,114] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
382. k1_xboole_0 = k3_subset_1(X4,X5) | ~m1_subset_1(X5,k1_zfmisc_1(X4)) | ~r1_tarski(X4,X5) [superposition 126,130] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
722. k1_xboole_0 != k1_struct_0(sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(sK1)) | ~r1_tarski(sK1,sK1) <- (~28) [superposition 378,382] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
735. 49 <=> k1_xboole_0 = k1_struct_0(sK0) [avatar definition]
736. k1_xboole_0 = k1_struct_0(sK0) <- (49) [avatar component clause 735] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
737. k1_xboole_0 != k1_struct_0(sK0) <- (~49) [avatar component clause 735] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
743. k1_xboole_0 != k1_xboole_0 | ~l1_struct_0(sK0) <- (~49) [superposition 737,119] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
744. ~l1_struct_0(sK0) <- (~49) [trivial inequality removal 743] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
749. $false <- (12, ~49) [subsumption resolution 744,203] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
750. ~12 | 49 [avatar contradiction clause 749]
751. k1_xboole_0 != k1_struct_0(sK0) | k1_xboole_0 != k3_subset_1(u1_struct_0(sK0),sK1) | k1_struct_0(sK0) = k3_subset_1(u1_struct_0(sK0),sK1) [theory tautology sat conflict] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:1,thDist:-1}
752. k1_xboole_0 != k1_struct_0(sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(sK1)) <- (~28) [subsumption resolution 722,136] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
753. ~21 | ~49 | 28 [avatar split clause 752,376,735,325]
754. u1_struct_0(sK0) != sK1 | ~l1_struct_0(sK0) <- (~2) [superposition 149,120] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
755. u1_struct_0(sK0) != sK1 <- (~2, 12) [subsumption resolution 754,203] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
756. ~4 | 2 | ~12 [avatar split clause 755,201,147,155]
759. ~r1_tarski(sK1,sK1) <- (~21) [resolution 326,128] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
761. $false <- (~21) [subsumption resolution 759,136] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
762. 21 [avatar contradiction clause 761]
826. k1_xboole_0 = k3_subset_1(u1_struct_0(sK0),sK1) <- (1, 49) [backward demodulation 144,736] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
838. 3 | ~1 | ~49 [avatar split clause 826,735,143,151]
858. sK1 = k3_subset_1(u1_struct_0(sK0),k1_xboole_0) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (3) [superposition 139,152] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
859. m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (3) [superposition 138,152] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
860. sK1 = k3_subset_1(u1_struct_0(sK0),k1_xboole_0) <- (3, 5) [subsumption resolution 858,165] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
862. 56 <=> sK1 = k3_subset_1(u1_struct_0(sK0),k1_xboole_0) [avatar definition]
864. sK1 = k3_subset_1(u1_struct_0(sK0),k1_xboole_0) <- (56) [avatar component clause 862] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
865. 56 | ~3 | ~5 [avatar split clause 860,163,151,862]
866. m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) <- (3, 5) [subsumption resolution 859,165] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
868. 57 <=> m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
870. m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) <- (57) [avatar component clause 868] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
871. 57 | ~3 | ~5 [avatar split clause 866,163,151,868]
891. u1_struct_0(sK0) = sK1 | ~m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) <- (56) [superposition 381,864] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
902. ~m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) <- (~4, 56) [subsumption resolution 891,157] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
903. $false <- (~4, 56, 57) [subsumption resolution 902,870] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
904. 4 | ~56 | ~57 [avatar contradiction clause 903]
907. $false [avatar sat refutation 158,161,166,171,204,314,379,750,751,753,756,762,838,865,871,904] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:586,thDist:-586}
% SZS output end Proof for tops_3__t2_tops_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 39
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 3
% 
% >>> Saturation
% Initial clauses: 49
% Generated clauses: 579
% Active clauses: 136
% Passive clauses: 217
% Final active clauses: 108
% Final passive clauses: 69
% Main loop iterations started: 148
% 
% >>> Simplifying Inferences
% Duplicate literals: 31
% Trivial inequalities: 2
% Fw subsumption resolutions: 52
% Fw demodulations: 15
% Bw demodulations: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 23
% Equational tautologies: 2
% Forward subsumptions: 160
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 165
% Unit resulting resolution: 8
% Forward superposition: 82
% Backward superposition: 88
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 20
% Split components: 48
% Sat splitting refutations: 13
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 87
% SAT solver unit clauses: 13
% SAT solver binary clauses: 18
% 
% Memory used [KB]: 9850
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% resolution: 0.001 s
% superposition: 0.001 s
% other: 0.006 s ( own 0.003 s ) 

