i: [1,54,0,2,1,1,0,0] t5_cgames_1
a: [3,111,0,2,1,54]
s: 111
i: [1,53,0,2,1,1,0,0] t5_cgames_1
a: [3,116,0,2,1,53]
s: 116
i: [1,57,0,13,3,1,0,0] t5_cgames_1
a: [3,96,0,3,1,57]
s: 96
a: [3,99,0,5,1,57]
s: 99
a: [3,103,0,5,1,57]
s: 103
i: [1,74,0,2,1,0,0,1] rc4_cgames_1
a: [3,121,0,2,1,74]
s: 121
i: [1,75,0,2,1,0,0,2] existence_l1_cgames_1
a: [3,126,0,2,1,75]
s: 126
i: [1,77,0,2,1,0,0,2] rc1_cgames_1
a: [3,131,0,2,1,77]
s: 131
i: [1,76,0,2,1,0,0,2] rc1_cgames_1
a: [3,136,0,2,1,76]
s: 136
i: [1,79,0,2,1,0,0,2] rc2_cgames_1
a: [3,141,0,2,1,79]
s: 141
i: [1,78,0,2,1,0,0,2] rc2_cgames_1
a: [3,146,0,2,1,78]
s: 146
i: [1,82,0,2,1,0,0,1] rc3_cgames_1
a: [3,151,0,2,1,82]
s: 151
i: [1,81,0,2,1,0,0,1] rc3_cgames_1
a: [3,156,0,2,1,81]
s: 156
i: [1,80,0,2,1,0,0,1] rc3_cgames_1
a: [3,161,0,2,1,80]
s: 161
i: [1,58,0,4,2,0,0,1] t4_cgames_1
s: 58
i: [1,59,0,4,2,0,0,1] t4_cgames_1
s: 59
d: [2,163,1,2,1,1,40,58,116]
a: [3,170,1,2,1,163]
s: 170
i: [1,70,0,4,1,0,0,3] dt_g1_cgames_1
s: 70
d: [2,164,1,2,1,1,40,58,111]
a: [3,175,1,2,1,164]
s: 175
i: [1,71,0,4,1,0,0,3] dt_g1_cgames_1
s: 71
i: [1,68,0,6,3,0,0,1] cc2_cgames_1
s: 68
d: [2,182,1,2,1,1,40,59,116]
a: [3,189,1,2,1,182]
s: 189
i: [1,87,0,7,2,0,0,1] d6_cgames_1
s: 87
d: [2,183,1,2,1,1,40,59,111]
a: [3,194,1,2,1,183]
s: 194
i: [1,92,0,7,2,0,0,1] d7_cgames_1
s: 92
i: [1,86,0,8,2,0,0,1] d6_cgames_1
s: 86
d: [2,201,1,7,2,1,44,87,99]
d: [2,202,1,5,1,2,21,201,111]
a: [3,206,1,5,1,202]
s: 206
i: [1,91,0,8,2,0,0,1] d7_cgames_1
s: 91
d: [2,208,1,7,2,1,44,92,103]
d: [2,209,1,5,1,2,21,208,111]
a: [3,213,1,5,1,209]
s: 213
i: [1,84,0,9,3,0,0,1] d6_cgames_1
s: 84
d: [2,235,0,7,2,0,21,84,59]
s: 235
d: [2,215,1,8,2,1,44,86,99]
d: [2,218,1,6,1,2,21,215,111]
a: [3,222,1,6,1,218]
s: 222
r: 206
i: [1,89,0,9,3,0,0,1] d7_cgames_1
s: 89
d: [2,224,1,8,2,1,44,91,103]
d: [2,227,1,6,1,2,21,224,111]
a: [3,231,1,6,1,227]
s: 231
r: 213
d: [2,237,0,7,2,0,21,89,59]
s: 237
i: [1,72,0,10,2,0,0,3] free_g1_cgames_1
s: 72
d: [2,165,1,2,1,1,40,58,121]
a: [3,180,1,2,1,165]
s: 180
i: [1,73,0,10,2,0,0,3] free_g1_cgames_1
s: 73
d: [2,184,1,2,1,1,40,59,121]
a: [3,199,1,2,1,184]
s: 199
i: [1,85,0,10,2,0,0,1] d6_cgames_1
s: 85
i: [1,90,0,10,2,0,0,1] d7_cgames_1
s: 90
d: [2,241,1,10,2,1,44,85,99]
d: [2,246,1,8,1,2,21,241,111]
d: [2,247,1,5,1,3,22,246,222]
a: [3,251,1,5,1,247]
s: 251
i: [1,69,0,11,3,0,0,2] abstractness_v1_cgames_1
s: 69
d: [2,253,1,10,2,1,44,90,103]
d: [2,258,1,8,1,2,21,253,111]
d: [2,259,1,5,1,3,22,258,231]
a: [3,263,1,5,1,259]
s: 263
d: [2,267,1,13,4,0,44,69,235]
d: [2,278,1,11,3,0,21,267,58]
d: [2,279,1,9,2,0,21,278,59]
s: 279
d: [2,270,1,13,4,0,44,69,237]
d: [2,288,1,11,3,0,21,270,58]
d: [2,289,1,9,2,0,21,288,59]
s: 289
d: [2,269,2,11,3,1,44,69,251]
d: [2,281,2,9,2,2,21,269,175]
d: [2,282,2,7,1,3,21,281,194]
a: [3,286,2,7,1,282]
s: 286
d: [2,272,1,13,4,0,44,73,69]
s: 272
d: [2,292,2,11,3,1,44,69,263]
d: [2,294,2,9,2,2,21,292,175]
d: [2,295,2,7,1,3,21,294,194]
d: [2,296,2,7,1,4,22,295,251]
a: [3,300,2,7,1,296]
s: 300
d: [2,273,1,13,4,0,44,73,69]
s: 273
d: [2,274,1,13,4,0,44,72,69]
s: 274
d: [2,346,3,9,2,1,44,72,300]
s: 346
d: [2,275,1,13,4,0,44,72,69]
s: 275
d: [2,347,3,9,2,1,44,72,300]
s: 347
d: [2,243,1,15,3,0,44,85,235]
s: 243
d: [2,244,1,15,3,0,44,235,85]
s: 244
d: [2,348,3,9,2,1,44,73,300]
s: 348
d: [2,255,1,15,3,0,44,90,237]
s: 255
d: [2,349,3,9,2,1,44,73,300]
s: 349
d: [2,256,1,15,3,0,44,237,90]
s: 256
d: [2,268,1,25,4,0,44,69,85]
d: [2,280,1,20,3,0,21,268,87]
s: 280
d: [2,359,4,10,3,1,44,346,279]
s: 359
d: [2,271,1,25,4,0,44,69,90]
d: [2,290,1,20,3,0,21,271,92]
s: 290
d: [2,360,4,10,3,1,44,346,289]
s: 360
d: [2,341,2,14,3,0,47,272]
d: [2,342,2,10,2,0,21,341,70]
d: [2,343,2,6,1,0,21,342,71]
s: 343
d: [2,355,2,14,3,0,47,274]
d: [2,356,2,10,2,0,21,355,70]
d: [2,357,2,6,1,0,21,356,71]
s: 357
d: [2,376,4,10,3,1,44,348,279]
s: 376
d: [2,303,2,11,3,0,44,279,237]
d: [2,312,2,9,2,0,18,303]
s: 312
d: [2,457,3,5,2,1,44,312,300]
d: [2,477,3,2,1,2,21,457,96]
d: [2,478,3,0,0,3,21,477,116]
e: 478
d: [2,458,3,5,2,1,44,300,312]
d: [2,480,3,2,1,2,21,458,96]
d: [2,481,3,0,0,3,21,480,116]
e: 481
i: [1,56,0,8,2,1,0,0] t5_cgames_1
i: [1,55,0,8,2,1,0,0] t5_cgames_1
d: [2,305,2,9,2,1,44,279,263]
d: [2,316,2,7,1,2,21,305,111]
d: [2,317,2,7,1,3,22,316,99]
f: 481,317,259,227,53,54,55,56,57
% Refutation found. Thanks to Tanya!
% SZS status Theorem for cgames_1__t5_cgames_1
% SZS output start Proof for cgames_1__t5_cgames_1
1. ! [X0] : (v2_cgames_1(X0) => ! [X1] : (v2_cgames_1(X1) => (X0 = X1 <=> (k7_cgames_1(X0) = k7_cgames_1(X1) & k6_cgames_1(X0) = k6_cgames_1(X1))))) [input t5_cgames_1]
2. ~! [X0] : (v2_cgames_1(X0) => ! [X1] : (v2_cgames_1(X1) => (X0 = X1 <=> (k7_cgames_1(X0) = k7_cgames_1(X1) & k6_cgames_1(X0) = k6_cgames_1(X1))))) [negated conjecture 1]
3. ! [X0] : (l1_cgames_1(X0) => (v1_cgames_1(X0) => g1_cgames_1(u1_cgames_1(X0),u2_cgames_1(X0)) = X0)) [input abstractness_v1_cgames_1]
5. ! [X0] : (v2_cgames_1(X0) => ! [X1] : (k6_cgames_1(X0) = X1 <=> ? [X2] : (u1_cgames_1(X2) = X1 & X0 = X2 & l1_cgames_1(X2)))) [input d6_cgames_1]
6. ! [X0] : (v2_cgames_1(X0) => ! [X1] : (k7_cgames_1(X0) = X1 <=> ? [X2] : (u2_cgames_1(X2) = X1 & X0 = X2 & l1_cgames_1(X2)))) [input d7_cgames_1]
19. ! [X0] : (v2_cgames_1(X0) => (l1_cgames_1(X0) & v1_cgames_1(X0))) [input t4_cgames_1]
21. ? [X0] : (? [X1] : ((X0 = X1 <~> (k7_cgames_1(X0) = k7_cgames_1(X1) & k6_cgames_1(X0) = k6_cgames_1(X1))) & v2_cgames_1(X1)) & v2_cgames_1(X0)) [ennf transformation 2]
22. ! [X0] : ((l1_cgames_1(X0) & v1_cgames_1(X0)) | ~v2_cgames_1(X0)) [ennf transformation 19]
23. ! [X0] : (! [X1] : (k6_cgames_1(X0) = X1 <=> ? [X2] : (u1_cgames_1(X2) = X1 & X0 = X2 & l1_cgames_1(X2))) | ~v2_cgames_1(X0)) [ennf transformation 5]
24. ! [X0] : (! [X1] : (k7_cgames_1(X0) = X1 <=> ? [X2] : (u2_cgames_1(X2) = X1 & X0 = X2 & l1_cgames_1(X2))) | ~v2_cgames_1(X0)) [ennf transformation 6]
27. ! [X0] : ((g1_cgames_1(u1_cgames_1(X0),u2_cgames_1(X0)) = X0 | ~v1_cgames_1(X0)) | ~l1_cgames_1(X0)) [ennf transformation 3]
28. ! [X0] : (g1_cgames_1(u1_cgames_1(X0),u2_cgames_1(X0)) = X0 | ~v1_cgames_1(X0) | ~l1_cgames_1(X0)) [flattening 27]
30. ? [X0] : (? [X1] : ((((k7_cgames_1(X0) != k7_cgames_1(X1) | k6_cgames_1(X0) != k6_cgames_1(X1)) | X0 != X1) & ((k7_cgames_1(X0) = k7_cgames_1(X1) & k6_cgames_1(X0) = k6_cgames_1(X1)) | X0 = X1)) & v2_cgames_1(X1)) & v2_cgames_1(X0)) [nnf transformation 21]
31. ? [X0] : (? [X1] : ((k7_cgames_1(X0) != k7_cgames_1(X1) | k6_cgames_1(X0) != k6_cgames_1(X1) | X0 != X1) & ((k7_cgames_1(X0) = k7_cgames_1(X1) & k6_cgames_1(X0) = k6_cgames_1(X1)) | X0 = X1) & v2_cgames_1(X1)) & v2_cgames_1(X0)) [flattening 30]
32. ? [X0] : (? [X1] : ((k7_cgames_1(X0) != k7_cgames_1(X1) | k6_cgames_1(X0) != k6_cgames_1(X1) | X0 != X1) & ((k7_cgames_1(X0) = k7_cgames_1(X1) & k6_cgames_1(X0) = k6_cgames_1(X1)) | X0 = X1) & v2_cgames_1(X1)) & v2_cgames_1(X0)) => (? [X1] : ((k7_cgames_1(X1) != k7_cgames_1(sK0) | k6_cgames_1(X1) != k6_cgames_1(sK0) | sK0 != X1) & ((k7_cgames_1(X1) = k7_cgames_1(sK0) & k6_cgames_1(X1) = k6_cgames_1(sK0)) | sK0 = X1) & v2_cgames_1(X1)) & v2_cgames_1(sK0)) [choice axiom]
33. ? [X1] : ((k7_cgames_1(X1) != k7_cgames_1(sK0) | k6_cgames_1(X1) != k6_cgames_1(sK0) | sK0 != X1) & ((k7_cgames_1(X1) = k7_cgames_1(sK0) & k6_cgames_1(X1) = k6_cgames_1(sK0)) | sK0 = X1) & v2_cgames_1(X1)) => ((k7_cgames_1(sK0) != k7_cgames_1(sK1) | k6_cgames_1(sK0) != k6_cgames_1(sK1) | sK0 != sK1) & ((k7_cgames_1(sK0) = k7_cgames_1(sK1) & k6_cgames_1(sK0) = k6_cgames_1(sK1)) | sK0 = sK1) & v2_cgames_1(sK1)) [choice axiom]
34. ((k7_cgames_1(sK0) != k7_cgames_1(sK1) | k6_cgames_1(sK0) != k6_cgames_1(sK1) | sK0 != sK1) & ((k7_cgames_1(sK0) = k7_cgames_1(sK1) & k6_cgames_1(sK0) = k6_cgames_1(sK1)) | sK0 = sK1) & v2_cgames_1(sK1)) & v2_cgames_1(sK0) [skolemisation 31,33,32]
35. ! [X0] : (! [X1] : ((k6_cgames_1(X0) = X1 | ! [X2] : (u1_cgames_1(X2) != X1 | X0 != X2 | ~l1_cgames_1(X2))) & (? [X2] : (u1_cgames_1(X2) = X1 & X0 = X2 & l1_cgames_1(X2)) | k6_cgames_1(X0) != X1)) | ~v2_cgames_1(X0)) [nnf transformation 23]
36. ! [X0] : (! [X1] : ((k6_cgames_1(X0) = X1 | ! [X2] : (u1_cgames_1(X2) != X1 | X0 != X2 | ~l1_cgames_1(X2))) & (? [X3] : (u1_cgames_1(X3) = X1 & X0 = X3 & l1_cgames_1(X3)) | k6_cgames_1(X0) != X1)) | ~v2_cgames_1(X0)) [rectify 35]
37. ! [X1,X0] : (? [X3] : (u1_cgames_1(X3) = X1 & X0 = X3 & l1_cgames_1(X3)) => (u1_cgames_1(sK2(X0,X1)) = X1 & sK2(X0,X1) = X0 & l1_cgames_1(sK2(X0,X1)))) [choice axiom]
38. ! [X0] : (! [X1] : ((k6_cgames_1(X0) = X1 | ! [X2] : (u1_cgames_1(X2) != X1 | X0 != X2 | ~l1_cgames_1(X2))) & ((u1_cgames_1(sK2(X0,X1)) = X1 & sK2(X0,X1) = X0 & l1_cgames_1(sK2(X0,X1))) | k6_cgames_1(X0) != X1)) | ~v2_cgames_1(X0)) [skolemisation 36,37]
39. ! [X0] : (! [X1] : ((k7_cgames_1(X0) = X1 | ! [X2] : (u2_cgames_1(X2) != X1 | X0 != X2 | ~l1_cgames_1(X2))) & (? [X2] : (u2_cgames_1(X2) = X1 & X0 = X2 & l1_cgames_1(X2)) | k7_cgames_1(X0) != X1)) | ~v2_cgames_1(X0)) [nnf transformation 24]
40. ! [X0] : (! [X1] : ((k7_cgames_1(X0) = X1 | ! [X2] : (u2_cgames_1(X2) != X1 | X0 != X2 | ~l1_cgames_1(X2))) & (? [X3] : (u2_cgames_1(X3) = X1 & X0 = X3 & l1_cgames_1(X3)) | k7_cgames_1(X0) != X1)) | ~v2_cgames_1(X0)) [rectify 39]
41. ! [X1,X0] : (? [X3] : (u2_cgames_1(X3) = X1 & X0 = X3 & l1_cgames_1(X3)) => (u2_cgames_1(sK3(X0,X1)) = X1 & sK3(X0,X1) = X0 & l1_cgames_1(sK3(X0,X1)))) [choice axiom]
42. ! [X0] : (! [X1] : ((k7_cgames_1(X0) = X1 | ! [X2] : (u2_cgames_1(X2) != X1 | X0 != X2 | ~l1_cgames_1(X2))) & ((u2_cgames_1(sK3(X0,X1)) = X1 & sK3(X0,X1) = X0 & l1_cgames_1(sK3(X0,X1))) | k7_cgames_1(X0) != X1)) | ~v2_cgames_1(X0)) [skolemisation 40,41]
53. v2_cgames_1(sK0) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. v2_cgames_1(sK1) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. k6_cgames_1(sK0) = k6_cgames_1(sK1) | sK0 = sK1 [cnf transformation 34] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. k7_cgames_1(sK0) = k7_cgames_1(sK1) | sK0 = sK1 [cnf transformation 34] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
57. k7_cgames_1(sK0) != k7_cgames_1(sK1) | k6_cgames_1(sK0) != k6_cgames_1(sK1) | sK0 != sK1 [cnf transformation 34] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. ~v2_cgames_1(X0) | v1_cgames_1(X0) [cnf transformation 22] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. ~v2_cgames_1(X0) | l1_cgames_1(X0) [cnf transformation 22] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. k6_cgames_1(X0) = X1 | u1_cgames_1(X2) != X1 | X0 != X2 | ~l1_cgames_1(X2) | ~v2_cgames_1(X0) [cnf transformation 38] {a:0,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
65. sK3(X0,X1) = X0 | k7_cgames_1(X0) != X1 | ~v2_cgames_1(X0) [cnf transformation 42] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
66. u2_cgames_1(sK3(X0,X1)) = X1 | k7_cgames_1(X0) != X1 | ~v2_cgames_1(X0) [cnf transformation 42] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
67. k7_cgames_1(X0) = X1 | u2_cgames_1(X2) != X1 | X0 != X2 | ~l1_cgames_1(X2) | ~v2_cgames_1(X0) [cnf transformation 42] {a:0,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
69. g1_cgames_1(u1_cgames_1(X0),u2_cgames_1(X0)) = X0 | ~v1_cgames_1(X0) | ~l1_cgames_1(X0) [cnf transformation 28] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
83. k6_cgames_1(X0) = u1_cgames_1(X2) | X0 != X2 | ~l1_cgames_1(X2) | ~v2_cgames_1(X0) [equality resolution 63] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
84. u1_cgames_1(X2) = k6_cgames_1(X2) | ~l1_cgames_1(X2) | ~v2_cgames_1(X2) [equality resolution 83] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
88. k7_cgames_1(X0) = u2_cgames_1(X2) | X0 != X2 | ~l1_cgames_1(X2) | ~v2_cgames_1(X0) [equality resolution 67] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
89. u2_cgames_1(X2) = k7_cgames_1(X2) | ~l1_cgames_1(X2) | ~v2_cgames_1(X2) [equality resolution 88] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
90. k7_cgames_1(X0) = u2_cgames_1(sK3(X0,k7_cgames_1(X0))) | ~v2_cgames_1(X0) [equality resolution 66] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
91. sK3(X0,k7_cgames_1(X0)) = X0 | ~v2_cgames_1(X0) [equality resolution 65] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
94. 1 <=> sK0 = sK1 [avatar definition]
96. sK0 != sK1 <- (~1) [avatar component clause 94] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. 2 <=> k6_cgames_1(sK0) = k6_cgames_1(sK1) [avatar definition]
99. k6_cgames_1(sK0) = k6_cgames_1(sK1) <- (2) [avatar component clause 98] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
102. 3 <=> k7_cgames_1(sK0) = k7_cgames_1(sK1) [avatar definition]
103. k7_cgames_1(sK0) = k7_cgames_1(sK1) <- (3) [avatar component clause 102] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. ~1 | ~2 | ~3 [avatar split clause 57,102,98,94]
106. 1 | 3 [avatar split clause 56,102,94]
107. 1 | 2 [avatar split clause 55,98,94]
109. 4 <=> v2_cgames_1(sK1) [avatar definition]
111. v2_cgames_1(sK1) <- (4) [avatar component clause 109] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
112. 4 [avatar split clause 54,109]
114. 5 <=> v2_cgames_1(sK0) [avatar definition]
116. v2_cgames_1(sK0) <- (5) [avatar component clause 114] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
117. 5 [avatar split clause 53,114]
224. sK1 = sK3(sK1,k7_cgames_1(sK0)) | ~v2_cgames_1(sK1) <- (3) [superposition 91,103] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
227. sK1 = sK3(sK1,k7_cgames_1(sK0)) <- (3, 4) [subsumption resolution 224,111] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
229. 24 <=> sK1 = sK3(sK1,k7_cgames_1(sK0)) [avatar definition]
231. sK1 = sK3(sK1,k7_cgames_1(sK0)) <- (24) [avatar component clause 229] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
232. 24 | ~3 | ~4 [avatar split clause 227,109,102,229]
233. sK0 != sK1 | k7_cgames_1(sK0) = k7_cgames_1(sK1) [theory tautology sat conflict] {a:0,w:8,wCS:3200,sine:255,thAx:0,allAx:1,thDist:-1}
234. sK0 != sK1 | k6_cgames_1(sK0) = k6_cgames_1(sK1) [theory tautology sat conflict] {a:0,w:8,wCS:3200,sine:255,thAx:0,allAx:1,thDist:-1}
235. u1_cgames_1(X2) = k6_cgames_1(X2) | ~v2_cgames_1(X2) [subsumption resolution 84,59] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
237. u2_cgames_1(X2) = k7_cgames_1(X2) | ~v2_cgames_1(X2) [subsumption resolution 89,59] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
253. k7_cgames_1(sK0) = u2_cgames_1(sK3(sK1,k7_cgames_1(sK0))) | ~v2_cgames_1(sK1) <- (3) [superposition 90,103] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
258. k7_cgames_1(sK0) = u2_cgames_1(sK3(sK1,k7_cgames_1(sK0))) <- (3, 4) [subsumption resolution 253,111] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
259. k7_cgames_1(sK0) = u2_cgames_1(sK1) <- (3, 4, 24) [forward demodulation 258,231] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
261. 26 <=> k7_cgames_1(sK0) = u2_cgames_1(sK1) [avatar definition]
263. k7_cgames_1(sK0) = u2_cgames_1(sK1) <- (26) [avatar component clause 261] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
264. 26 | ~3 | ~4 | ~24 [avatar split clause 259,229,109,102,261]
267. g1_cgames_1(k6_cgames_1(X0),u2_cgames_1(X0)) = X0 | ~v1_cgames_1(X0) | ~l1_cgames_1(X0) | ~v2_cgames_1(X0) [superposition 69,235] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:3,thDist:-3}
278. g1_cgames_1(k6_cgames_1(X0),u2_cgames_1(X0)) = X0 | ~l1_cgames_1(X0) | ~v2_cgames_1(X0) [subsumption resolution 267,58] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:3,thDist:-3}
279. g1_cgames_1(k6_cgames_1(X0),u2_cgames_1(X0)) = X0 | ~v2_cgames_1(X0) [subsumption resolution 278,59] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
298. 28 <=> sK1 = g1_cgames_1(k6_cgames_1(sK0),k7_cgames_1(sK0)) [avatar definition]
300. sK1 = g1_cgames_1(k6_cgames_1(sK0),k7_cgames_1(sK0)) <- (28) [avatar component clause 298] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
303. g1_cgames_1(k6_cgames_1(X0),k7_cgames_1(X0)) = X0 | ~v2_cgames_1(X0) | ~v2_cgames_1(X0) [superposition 279,237] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:5,thDist:-5}
305. sK1 = g1_cgames_1(k6_cgames_1(sK1),k7_cgames_1(sK0)) | ~v2_cgames_1(sK1) <- (26) [superposition 279,263] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
312. g1_cgames_1(k6_cgames_1(X0),k7_cgames_1(X0)) = X0 | ~v2_cgames_1(X0) [duplicate literal removal 303] {a:2,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
316. sK1 = g1_cgames_1(k6_cgames_1(sK1),k7_cgames_1(sK0)) <- (4, 26) [subsumption resolution 305,111] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
317. sK1 = g1_cgames_1(k6_cgames_1(sK0),k7_cgames_1(sK0)) <- (2, 4, 26) [forward demodulation 316,99] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
318. 28 | ~2 | ~4 | ~26 [avatar split clause 317,261,109,98,298]
458. sK0 = sK1 | ~v2_cgames_1(sK0) <- (28) [superposition 300,312] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
480. ~v2_cgames_1(sK0) <- (~1, 28) [subsumption resolution 458,96] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
481. $false <- (~1, 5, 28) [subsumption resolution 480,116] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
482. 1 | ~5 | ~28 [avatar contradiction clause 481]
483. $false [avatar sat refutation 105,106,107,112,117,232,233,234,264,318,482] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:233,thDist:-233}
% SZS output end Proof for cgames_1__t5_cgames_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 30
% Generated clauses: 297
% Active clauses: 58
% Passive clauses: 129
% Final active clauses: 56
% Final passive clauses: 61
% Main loop iterations started: 68
% 
% >>> Simplifying Inferences
% Duplicate literals: 22
% Fw subsumption resolutions: 42
% Fw demodulations: 16
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Equational tautologies: 2
% Forward subsumptions: 49
% Fw demodulations to eq. taut.: 1
% 
% >>> Generating Inferences
% Binary resolution: 8
% Forward superposition: 61
% Backward superposition: 79
% Self superposition: 2
% Equality resolution: 7
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 39
% SAT solver unit clauses: 11
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.002 s ( own 0.0000 s ) 

