i: [1,28,0,2,1,1,0,0] t18_orders_1
a: [3,52,0,2,1,28]
s: 52
i: [1,27,0,2,1,1,0,0] t18_orders_1
a: [3,57,0,2,1,27]
s: 57
i: [1,29,0,3,1,1,0,0] t18_orders_1
a: [3,47,0,3,1,29]
s: 47
i: [1,30,0,5,2,0,0,1] dt_k2_relat_1
s: 30
i: [1,33,0,6,3,0,0,1] d5_orders_1
s: 33
i: [1,34,0,6,3,0,0,1] d5_orders_1
s: 34
d: [2,59,1,4,2,1,40,33,52]
d: [2,60,1,2,1,2,21,59,57]
a: [3,64,1,2,1,60]
s: 64
i: [1,35,0,6,3,0,0,1] d5_orders_1
s: 35
d: [2,66,1,4,2,1,40,34,52]
d: [2,67,1,2,1,2,21,66,57]
a: [3,71,1,2,1,67]
s: 71
i: [1,36,0,6,3,0,0,1] d5_orders_1
s: 36
i: [1,31,0,7,2,0,0,255] involutiveness_k2_relat_1
s: 31
d: [2,73,1,4,2,1,40,35,52]
d: [2,74,1,2,1,2,21,73,57]
a: [3,78,1,2,1,74]
s: 78
i: [1,32,0,7,3,0,0,2] l25_orders_1
s: 32
d: [2,80,1,4,2,1,40,36,52]
d: [2,81,1,2,1,2,21,80,57]
a: [3,85,1,2,1,81]
s: 85
i: [1,38,0,7,3,0,0,2] fc17_relat_2
s: 38
i: [1,39,0,7,3,0,0,2] fc17_relat_2
s: 39
i: [1,40,0,7,3,0,0,2] fc2_relat_2
s: 40
i: [1,41,0,7,3,0,0,2] fc2_relat_2
s: 41
i: [1,42,0,7,3,0,0,2] fc21_relat_2
s: 42
i: [1,43,0,7,3,0,0,2] fc21_relat_2
s: 43
i: [1,37,0,12,6,0,0,1] d5_orders_1
s: 37
d: [2,88,1,10,4,0,44,32,31]
d: [2,89,1,7,3,0,21,88,30]
s: 89
d: [2,90,1,10,4,0,44,39,31]
d: [2,91,1,7,3,0,21,90,30]
s: 91
d: [2,92,1,10,4,0,44,41,31]
d: [2,93,1,7,3,0,21,92,30]
s: 93
d: [2,94,1,10,4,0,44,43,31]
d: [2,95,1,7,3,0,21,94,30]
s: 95
d: [2,96,1,19,7,0,40,37,32]
d: [2,98,1,16,6,0,21,96,30]
s: 98
d: [2,112,2,16,7,0,44,98,31]
d: [2,114,2,13,6,0,21,112,30]
s: 114
d: [2,111,2,17,7,0,40,98,39]
d: [2,113,2,15,6,0,18,111]
s: 113
d: [2,118,3,16,7,0,40,113,43]
d: [2,120,3,14,6,0,18,118]
s: 120
d: [2,119,3,17,7,0,44,113,31]
d: [2,121,3,14,6,0,21,119,30]
s: 121
d: [2,122,4,15,7,0,40,120,41]
d: [2,124,4,13,6,0,18,122]
s: 124
d: [2,129,5,0,0,6,52,64,71,78,57,85,47,124]
e: 129
d: [2,130,5,10,5,1,40,124,47]
d: [2,137,5,8,4,2,21,130,85]
d: [2,138,5,6,3,3,21,137,57]
d: [2,139,5,4,2,4,21,138,78]
d: [2,140,5,2,1,5,21,139,71]
d: [2,141,5,0,0,6,21,140,64]
e: 141
f: 141,81,74,67,60,27,28,29
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t18_orders_1
% SZS output start Proof for orders_1__t18_orders_1
1. ! [X0] : (v1_relat_1(X0) => (v3_orders_1(X0) => v3_orders_1(k2_relat_1(X0)))) [input t18_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => (v3_orders_1(X0) => v3_orders_1(k2_relat_1(X0)))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => (v3_orders_1(X0) <=> (v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d5_orders_1]
4. ! [X0] : (v1_relat_1(X0) => v1_relat_1(k2_relat_1(X0))) [input dt_k2_relat_1]
5. ! [X0] : ((v4_relat_2(X0) & v1_relat_1(X0)) => (v4_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0)))) [input fc17_relat_2]
6. ! [X0] : ((v8_relat_2(X0) & v1_relat_1(X0)) => (v8_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0)))) [input fc21_relat_2]
7. ! [X0] : ((v1_relat_2(X0) & v1_relat_1(X0)) => (v1_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0)))) [input fc2_relat_2]
9. ! [X0] : (v1_relat_1(X0) => (v6_relat_2(X0) => v6_relat_2(k2_relat_1(X0)))) [input l25_orders_1]
10. ? [X0] : ((~v3_orders_1(k2_relat_1(X0)) & v3_orders_1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
11. ? [X0] : (~v3_orders_1(k2_relat_1(X0)) & v3_orders_1(X0) & v1_relat_1(X0)) [flattening 10]
12. ! [X0] : (v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 4]
14. ! [X0] : ((v6_relat_2(k2_relat_1(X0)) | ~v6_relat_2(X0)) | ~v1_relat_1(X0)) [ennf transformation 9]
15. ! [X0] : (v6_relat_2(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0)) [flattening 14]
16. ! [X0] : ((v3_orders_1(X0) <=> (v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
17. ! [X0] : ((v4_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | (~v4_relat_2(X0) | ~v1_relat_1(X0))) [ennf transformation 5]
18. ! [X0] : ((v4_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | ~v4_relat_2(X0) | ~v1_relat_1(X0)) [flattening 17]
19. ! [X0] : ((v1_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | (~v1_relat_2(X0) | ~v1_relat_1(X0))) [ennf transformation 7]
20. ! [X0] : ((v1_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | ~v1_relat_2(X0) | ~v1_relat_1(X0)) [flattening 19]
21. ! [X0] : ((v8_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | (~v8_relat_2(X0) | ~v1_relat_1(X0))) [ennf transformation 6]
22. ! [X0] : ((v8_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | ~v8_relat_2(X0) | ~v1_relat_1(X0)) [flattening 21]
23. ? [X0] : (~v3_orders_1(k2_relat_1(X0)) & v3_orders_1(X0) & v1_relat_1(X0)) => (~v3_orders_1(k2_relat_1(sK0)) & v3_orders_1(sK0) & v1_relat_1(sK0)) [choice axiom]
24. ~v3_orders_1(k2_relat_1(sK0)) & v3_orders_1(sK0) & v1_relat_1(sK0) [skolemisation 11,23]
25. ! [X0] : (((v3_orders_1(X0) | (~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v3_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 16]
26. ! [X0] : (((v3_orders_1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v3_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 25]
27. v1_relat_1(sK0) [cnf transformation 24] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. v3_orders_1(sK0) [cnf transformation 24] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~v3_orders_1(k2_relat_1(sK0)) [cnf transformation 24] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 12] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
32. v6_relat_2(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 15] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
33. ~v3_orders_1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. ~v3_orders_1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
35. ~v3_orders_1(X0) | v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
36. ~v3_orders_1(X0) | v6_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. ~v6_relat_2(X0) | v3_orders_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
39. v4_relat_2(k2_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
41. v1_relat_2(k2_relat_1(X0)) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 20] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
43. v8_relat_2(k2_relat_1(X0)) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
45. 1 <=> v3_orders_1(k2_relat_1(sK0)) [avatar definition]
47. ~v3_orders_1(k2_relat_1(sK0)) <- (~1) [avatar component clause 45] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
48. ~1 [avatar split clause 29,45]
50. 2 <=> v3_orders_1(sK0) [avatar definition]
52. v3_orders_1(sK0) <- (2) [avatar component clause 50] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
53. 2 [avatar split clause 28,50]
55. 3 <=> v1_relat_1(sK0) [avatar definition]
57. v1_relat_1(sK0) <- (3) [avatar component clause 55] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
58. 3 [avatar split clause 27,55]
59. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 33,52] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
60. v1_relat_2(sK0) <- (2, 3) [subsumption resolution 59,57] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
62. 4 <=> v1_relat_2(sK0) [avatar definition]
64. v1_relat_2(sK0) <- (4) [avatar component clause 62] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
65. 4 | ~2 | ~3 [avatar split clause 60,55,50,62]
66. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 34,52] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
67. v8_relat_2(sK0) <- (2, 3) [subsumption resolution 66,57] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
69. 5 <=> v8_relat_2(sK0) [avatar definition]
71. v8_relat_2(sK0) <- (5) [avatar component clause 69] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. 5 | ~2 | ~3 [avatar split clause 67,55,50,69]
73. v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 35,52] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
74. v4_relat_2(sK0) <- (2, 3) [subsumption resolution 73,57] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
76. 6 <=> v4_relat_2(sK0) [avatar definition]
78. v4_relat_2(sK0) <- (6) [avatar component clause 76] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
79. 6 | ~2 | ~3 [avatar split clause 74,55,50,76]
80. v6_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 36,52] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
81. v6_relat_2(sK0) <- (2, 3) [subsumption resolution 80,57] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
83. 7 <=> v6_relat_2(sK0) [avatar definition]
85. v6_relat_2(sK0) <- (7) [avatar component clause 83] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. 7 | ~2 | ~3 [avatar split clause 81,55,50,83]
96. v3_orders_1(k2_relat_1(X0)) | ~v4_relat_2(k2_relat_1(X0)) | ~v8_relat_2(k2_relat_1(X0)) | ~v1_relat_2(k2_relat_1(X0)) | ~v1_relat_1(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) [resolution 37,32] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:2,thDist:-2}
98. ~v4_relat_2(k2_relat_1(X0)) | v3_orders_1(k2_relat_1(X0)) | ~v8_relat_2(k2_relat_1(X0)) | ~v1_relat_2(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) [subsumption resolution 96,30] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
111. v3_orders_1(k2_relat_1(X0)) | ~v8_relat_2(k2_relat_1(X0)) | ~v1_relat_2(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [resolution 98,39] {a:2,w:17,wCS:6800,sine:1,thAx:0,allAx:3,thDist:-3}
113. ~v8_relat_2(k2_relat_1(X0)) | v3_orders_1(k2_relat_1(X0)) | ~v1_relat_2(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) [duplicate literal removal 111] {a:2,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
118. v3_orders_1(k2_relat_1(X0)) | ~v1_relat_2(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [resolution 113,43] {a:3,w:16,wCS:6400,sine:1,thAx:0,allAx:4,thDist:-4}
120. ~v1_relat_2(k2_relat_1(X0)) | v3_orders_1(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) [duplicate literal removal 118] {a:3,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
122. v3_orders_1(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [resolution 120,41] {a:4,w:15,wCS:6000,sine:1,thAx:0,allAx:5,thDist:-5}
124. v3_orders_1(k2_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) [duplicate literal removal 122] {a:4,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
130. ~v6_relat_2(sK0) | ~v1_relat_1(sK0) | ~v4_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1) [resolution 124,47] {a:5,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
137. ~v1_relat_1(sK0) | ~v4_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1, 7) [subsumption resolution 130,85] {a:5,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
138. ~v4_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1, 3, 7) [subsumption resolution 137,57] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
139. ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1, 3, 6, 7) [subsumption resolution 138,78] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
140. ~v1_relat_2(sK0) <- (~1, 3, 5, 6, 7) [subsumption resolution 139,71] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
141. $false <- (~1, 3, 4, 5, 6, 7) [subsumption resolution 140,64] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
142. 1 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar contradiction clause 141]
144. $false [avatar sat refutation 48,53,58,65,72,79,86,142] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:61,thDist:-61}
% SZS output end Proof for orders_1__t18_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 1
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 87
% Active clauses: 28
% Passive clauses: 33
% Final active clauses: 28
% Final passive clauses: 2
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 18
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 15
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 1
% Forward superposition: 14
% Backward superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% superposition: 0.001 s ( own 0.0000 s ) 
% other: 0.001 s ( own 0.0000 s ) 

