i: [1,94,0,2,1,1,0,0] t45_subset_1
a: [3,142,0,2,1,94]
s: 142
i: [1,93,0,4,1,1,0,0] t45_subset_1
a: [3,147,0,4,1,93]
s: 147
i: [1,96,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,152,0,2,1,96]
s: 152
i: [1,97,0,2,1,0,0,255] fc1_xboole_0
a: [3,157,0,2,1,97]
s: 157
i: [1,134,0,2,1,0,0,1] rc2_zfmisc_1
a: [3,167,0,2,1,134]
s: 167
i: [1,135,0,2,1,0,0,2] rc2_xboole_0
a: [3,172,0,2,1,135]
s: 172
i: [1,136,0,2,1,0,0,2] rc1_xboole_0
a: [3,177,0,2,1,136]
s: 177
i: [1,138,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,182,0,2,1,138]
s: 182
i: [1,137,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,187,0,2,1,137]
s: 187
i: [1,99,0,3,1,0,0,1] fc1_subset_1
s: 99
i: [1,130,0,3,1,0,0,1] rc2_subset_1
s: 130
i: [1,98,0,3,1,0,0,255] d2_xboole_0
a: [3,162,0,3,1,98]
s: 162
r: 152
i: [1,95,0,15,5,1,0,0] t45_subset_1
s: 95
i: [1,100,0,4,2,0,0,1] cc2_zfmisc_1
s: 100
i: [1,113,0,4,2,0,0,1] cc1_zfmisc_1
s: 113
i: [1,128,0,4,1,0,0,1] existence_m1_subset_1
s: 128
d: [2,190,1,2,1,1,40,100,157]
a: [3,196,1,2,1,190]
s: 196
i: [1,132,0,4,1,0,0,255] rc3_subset_1
s: 132
d: [2,192,1,2,1,1,40,100,177]
a: [3,201,1,2,1,192]
s: 201
i: [1,102,0,5,2,0,0,1] rc6_subset_1
s: 102
i: [1,106,0,5,2,0,0,1] rc1_subset_1
s: 106
d: [2,191,1,3,1,0,40,100,130]
s: 191
i: [1,108,0,5,2,0,0,1] rc5_subset_1
s: 108
i: [1,109,0,5,2,0,0,1] rc5_subset_1
s: 109
i: [1,114,0,5,2,0,0,255] t6_boole
s: 114
i: [1,127,0,5,2,0,0,1] t7_boole
s: 127
i: [1,129,0,5,1,0,0,1] rc2_subset_1
s: 129
i: [1,131,0,5,1,0,0,255] rc3_subset_1
s: 131
i: [1,111,0,6,2,0,0,255] rc4_subset_1
s: 111
i: [1,118,0,6,2,0,0,1] d10_zfmisc_1
s: 118
i: [1,119,0,6,2,0,0,1] d10_zfmisc_1
s: 119
d: [2,206,1,16,5,0,40,118,95]
d: [2,208,1,14,4,1,21,206,142]
a: [3,212,1,4,1,208]
s: 212
i: [1,133,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 133
d: [2,217,1,16,5,0,40,119,95]
d: [2,219,1,14,4,1,21,217,142]
a: [3,223,1,4,1,219]
s: 223
i: [1,101,0,7,2,0,0,1] rc6_subset_1
s: 101
i: [1,105,0,7,2,0,0,1] rc1_subset_1
s: 105
d: [2,228,1,6,2,0,40,133,118]
s: 228
i: [1,107,0,7,2,0,0,1] rc5_subset_1
s: 107
d: [2,229,1,6,2,0,40,133,119]
s: 229
i: [1,110,0,7,2,0,0,255] rc4_subset_1
s: 110
i: [1,120,0,7,2,0,0,1] d10_zfmisc_1
s: 120
i: [1,123,0,7,3,0,0,1] d1_subset_1
s: 123
i: [1,124,0,7,3,0,0,1] d1_subset_1
s: 124
d: [2,238,2,5,2,1,40,124,212]
a: [3,245,2,2,1,238]
s: 245
i: [1,126,0,7,3,0,0,1] t8_boole
s: 126
i: [1,112,0,8,3,0,0,1] cc5_subset_1
s: 112
d: [2,257,1,4,2,1,40,112,147]
d: [2,265,1,2,1,2,21,257,142]
a: [3,269,1,2,1,265]
s: 269
i: [1,115,0,8,3,0,0,1] cc1_subset_1
s: 115
d: [2,235,1,5,2,0,40,123,128]
s: 235
i: [1,121,0,8,3,0,0,1] d1_subset_1
s: 121
i: [1,122,0,8,3,0,0,1] d1_subset_1
s: 122
d: [2,282,1,7,2,1,40,121,147]
d: [2,290,1,4,1,1,21,282,99]
a: [3,294,1,4,1,290]
s: 294
d: [2,302,0,6,2,0,21,122,127]
s: 302
d: [2,303,2,4,1,1,40,294,133]
a: [3,308,2,4,1,303]
s: 308
i: [1,116,0,9,3,0,0,255] cc4_subset_1
s: 116
i: [1,125,0,10,3,0,0,1] l2_subset_1
s: 125
d: [2,323,1,6,2,1,40,125,147]
s: 323
i: [1,103,0,11,4,0,0,255] cc3_subset_1
s: 103
d: [2,346,1,7,3,1,40,103,147]
d: [2,358,1,5,2,2,21,346,245]
a: [3,366,1,2,1,358]
s: 366
i: [1,104,0,11,4,0,0,255] cc2_subset_1
s: 104
i: [1,117,0,11,4,0,0,1] d10_zfmisc_1
s: 117
d: [2,373,2,10,3,1,40,117,294]
a: [3,377,2,3,1,373]
s: 377
d: [2,264,1,5,2,0,40,112,131]
s: 264
d: [2,331,2,6,2,1,40,323,118]
d: [2,333,2,4,1,2,21,331,142]
a: [3,337,2,4,1,333]
s: 337
d: [2,382,3,0,0,2,52,212,337,302]
e: 382
d: [2,384,3,4,1,1,40,337,302]
d: [2,388,3,0,0,2,21,384,212]
e: 388
r: 212
a: [3,211,1,4,1,208]
s: 211
a: [3,215,1,10,3,208]
s: 215
d: [2,332,2,6,2,1,40,323,119]
d: [2,339,2,4,1,2,21,332,142]
a: [3,343,2,4,1,339]
s: 343
d: [2,403,3,0,0,2,52,223,343,302]
e: 403
d: [2,405,3,4,1,1,40,343,302]
d: [2,409,3,0,0,2,21,405,223]
e: 409
r: 223
a: [3,222,1,4,1,219]
s: 222
d: [2,385,3,4,1,1,40,337,133]
a: [3,394,3,4,1,385]
s: 394
a: [3,226,1,10,3,219]
s: 226
d: [2,279,1,5,2,0,40,115,131]
s: 279
d: [2,406,3,4,1,1,40,343,133]
a: [3,415,3,4,1,406]
s: 415
d: [2,288,1,8,2,0,40,121,129]
d: [2,300,1,5,1,0,21,288,99]
s: 300
d: [2,423,2,12,4,1,44,120,226]
d: [2,426,2,10,3,2,21,423,142]
d: [2,427,2,6,2,3,21,426,215]
s: 427
r: 95
r: 215
r: 226
d: [2,434,3,6,2,3,40,427,118]
d: [2,436,3,2,1,4,21,434,211]
d: [2,437,3,0,0,4,21,436,142]
e: 437
d: [2,435,3,6,2,3,40,427,119]
d: [2,439,3,2,1,4,21,435,222]
d: [2,440,3,0,0,4,21,439,142]
e: 440
f: 440,405,384,339,333,219,208,93,94
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t45_subset_1
% SZS output start Proof for subset_1__t45_subset_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ~(! [X2] : (m1_subset_1(X2,X0) => ! [X3] : (m1_subset_1(X3,X0) => ~(X2 != X3 & r2_hidden(X3,X1) & r2_hidden(X2,X1)))) & ~v1_zfmisc_1(X1))) [input t45_subset_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ~(! [X2] : (m1_subset_1(X2,X0) => ! [X3] : (m1_subset_1(X3,X0) => ~(X2 != X3 & r2_hidden(X3,X1) & r2_hidden(X2,X1)))) & ~v1_zfmisc_1(X1))) [negated conjecture 1]
11. ! [X0] : (v1_zfmisc_1(X0) <=> ! [X1] : ! [X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => X1 = X2)) [input d10_zfmisc_1]
12. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
21. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (r2_hidden(X2,X1) => r2_hidden(X2,X0))) [input l2_subset_1]
33. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
35. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ~(! [X2] : (m1_subset_1(X2,X0) => ! [X3] : (m1_subset_1(X3,X0) => ~(X2 != X3 & r2_hidden(X3,X1) & r2_hidden(X2,X1)))) & ~v1_zfmisc_1(X1))) [flattening 2]
36. ! [X0] : (v1_zfmisc_1(X0) <=> ! [X1,X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => X1 = X2)) [flattening 11]
37. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 12]
38. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (r2_hidden(X2,X1) => r2_hidden(X2,X0))) [flattening 21]
40. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 33]
41. ? [X0,X1] : ((! [X2] : (! [X3] : ((X2 = X3 | ~r2_hidden(X3,X1) | ~r2_hidden(X2,X1)) | ~m1_subset_1(X3,X0)) | ~m1_subset_1(X2,X0)) & ~v1_zfmisc_1(X1)) & m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 35]
42. ? [X0,X1] : (! [X2] : (! [X3] : (X2 = X3 | ~r2_hidden(X3,X1) | ~r2_hidden(X2,X1) | ~m1_subset_1(X3,X0)) | ~m1_subset_1(X2,X0)) & ~v1_zfmisc_1(X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 41]
57. ! [X0] : (v1_zfmisc_1(X0) <=> ! [X1,X2] : (X1 = X2 | (~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)))) [ennf transformation 36]
58. ! [X0] : (v1_zfmisc_1(X0) <=> ! [X1,X2] : (X1 = X2 | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0))) [flattening 57]
59. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 37]
60. ! [X0,X1] : (! [X2] : (r2_hidden(X2,X0) | ~r2_hidden(X2,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 38]
62. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 40]
64. ? [X0,X1] : (! [X2] : (! [X3] : (X2 = X3 | ~r2_hidden(X3,X1) | ~r2_hidden(X2,X1) | ~m1_subset_1(X3,X0)) | ~m1_subset_1(X2,X0)) & ~v1_zfmisc_1(X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (! [X2] : (! [X3] : (X2 = X3 | ~r2_hidden(X3,sK1) | ~r2_hidden(X2,sK1) | ~m1_subset_1(X3,sK0)) | ~m1_subset_1(X2,sK0)) & ~v1_zfmisc_1(sK1) & m1_subset_1(sK1,k1_zfmisc_1(sK0))) [choice axiom]
65. ! [X2] : (! [X3] : (X2 = X3 | ~r2_hidden(X3,sK1) | ~r2_hidden(X2,sK1) | ~m1_subset_1(X3,sK0)) | ~m1_subset_1(X2,sK0)) & ~v1_zfmisc_1(sK1) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) [skolemisation 42,64]
74. ! [X0] : ((v1_zfmisc_1(X0) | ? [X1,X2] : (X1 != X2 & r2_hidden(X2,X0) & r2_hidden(X1,X0))) & (! [X1,X2] : (X1 = X2 | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)) | ~v1_zfmisc_1(X0))) [nnf transformation 58]
75. ! [X0] : ((v1_zfmisc_1(X0) | ? [X1,X2] : (X1 != X2 & r2_hidden(X2,X0) & r2_hidden(X1,X0))) & (! [X3,X4] : (X3 = X4 | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | ~v1_zfmisc_1(X0))) [rectify 74]
76. ! [X0] : (? [X1,X2] : (X1 != X2 & r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (sK6(X0) != sK7(X0) & r2_hidden(sK7(X0),X0) & r2_hidden(sK6(X0),X0))) [choice axiom]
77. ! [X0] : ((v1_zfmisc_1(X0) | (sK6(X0) != sK7(X0) & r2_hidden(sK7(X0),X0) & r2_hidden(sK6(X0),X0))) & (! [X3,X4] : (X3 = X4 | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | ~v1_zfmisc_1(X0))) [skolemisation 75,76]
78. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 59]
93. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 65] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
94. ~v1_zfmisc_1(sK1) [cnf transformation 65] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
95. ~r2_hidden(X3,sK1) | X2 = X3 | ~r2_hidden(X2,sK1) | ~m1_subset_1(X3,sK0) | ~m1_subset_1(X2,sK0) [cnf transformation 65] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
118. r2_hidden(sK6(X0),X0) | v1_zfmisc_1(X0) [cnf transformation 77] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
119. r2_hidden(sK7(X0),X0) | v1_zfmisc_1(X0) [cnf transformation 77] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
120. sK6(X0) != sK7(X0) | v1_zfmisc_1(X0) [cnf transformation 77] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
122. m1_subset_1(X1,X0) | ~r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 78] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
125. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | ~r2_hidden(X2,X1) | r2_hidden(X2,X0) [cnf transformation 60] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
127. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 62] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
140. 1 <=> v1_zfmisc_1(sK1) [avatar definition]
142. ~v1_zfmisc_1(sK1) <- (~1) [avatar component clause 140] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
143. ~1 [avatar split clause 94,140]
145. 2 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
147. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (2) [avatar component clause 145] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
148. 2 [avatar split clause 93,145]
206. v1_zfmisc_1(sK1) | sK6(sK1) = X0 | ~r2_hidden(X0,sK1) | ~m1_subset_1(sK6(sK1),sK0) | ~m1_subset_1(X0,sK0) [resolution 118,95] {a:1,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
208. sK6(sK1) = X0 | ~r2_hidden(X0,sK1) | ~m1_subset_1(sK6(sK1),sK0) | ~m1_subset_1(X0,sK0) <- (~1) [subsumption resolution 206,142] {a:1,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
210. 13 <=> m1_subset_1(sK6(sK1),sK0) [avatar definition]
214. 14 <=> ! [X0] : (sK6(sK1) = X0 | ~m1_subset_1(X0,sK0) | ~r2_hidden(X0,sK1)) [avatar definition]
215. sK6(sK1) = X0 | ~m1_subset_1(X0,sK0) | ~r2_hidden(X0,sK1) <- (14) [avatar component clause 214] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
216. ~13 | 14 | 1 [avatar split clause 208,140,214,210]
217. v1_zfmisc_1(sK1) | sK7(sK1) = X0 | ~r2_hidden(X0,sK1) | ~m1_subset_1(sK7(sK1),sK0) | ~m1_subset_1(X0,sK0) [resolution 119,95] {a:1,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
219. sK7(sK1) = X0 | ~r2_hidden(X0,sK1) | ~m1_subset_1(sK7(sK1),sK0) | ~m1_subset_1(X0,sK0) <- (~1) [subsumption resolution 217,142] {a:1,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
221. 15 <=> m1_subset_1(sK7(sK1),sK0) [avatar definition]
222. m1_subset_1(sK7(sK1),sK0) <- (15) [avatar component clause 221] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
225. 16 <=> ! [X0] : (sK7(sK1) = X0 | ~m1_subset_1(X0,sK0) | ~r2_hidden(X0,sK1)) [avatar definition]
226. sK7(sK1) = X0 | ~m1_subset_1(X0,sK0) | ~r2_hidden(X0,sK1) <- (16) [avatar component clause 225] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
227. ~15 | 16 | 1 [avatar split clause 219,140,225,221]
302. ~r2_hidden(X1,X0) | m1_subset_1(X1,X0) [subsumption resolution 122,127] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
323. ~r2_hidden(X3,sK1) | r2_hidden(X3,sK0) <- (2) [resolution 125,147] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
331. r2_hidden(sK6(sK1),sK0) | v1_zfmisc_1(sK1) <- (2) [resolution 323,118] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
332. r2_hidden(sK7(sK1),sK0) | v1_zfmisc_1(sK1) <- (2) [resolution 323,119] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
333. r2_hidden(sK6(sK1),sK0) <- (~1, 2) [subsumption resolution 331,142] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
335. 23 <=> r2_hidden(sK6(sK1),sK0) [avatar definition]
337. r2_hidden(sK6(sK1),sK0) <- (23) [avatar component clause 335] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
338. 23 | 1 | ~2 [avatar split clause 333,145,140,335]
339. r2_hidden(sK7(sK1),sK0) <- (~1, 2) [subsumption resolution 332,142] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
341. 24 <=> r2_hidden(sK7(sK1),sK0) [avatar definition]
343. r2_hidden(sK7(sK1),sK0) <- (24) [avatar component clause 341] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
344. 24 | 1 | ~2 [avatar split clause 339,145,140,341]
384. m1_subset_1(sK6(sK1),sK0) <- (23) [resolution 337,302] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
390. 13 | ~23 [avatar split clause 384,335,210]
405. m1_subset_1(sK7(sK1),sK0) <- (24) [resolution 343,302] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
411. 15 | ~24 [avatar split clause 405,341,221]
423. sK6(sK1) != X2 | v1_zfmisc_1(sK1) | ~m1_subset_1(X2,sK0) | ~r2_hidden(X2,sK1) <- (16) [superposition 120,226] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
426. sK6(sK1) != X2 | ~m1_subset_1(X2,sK0) | ~r2_hidden(X2,sK1) <- (~1, 16) [subsumption resolution 423,142] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
427. ~r2_hidden(X2,sK1) | ~m1_subset_1(X2,sK0) <- (~1, 14, 16) [subsumption resolution 426,215] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
435. ~m1_subset_1(sK7(sK1),sK0) | v1_zfmisc_1(sK1) <- (~1, 14, 16) [resolution 427,119] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
439. v1_zfmisc_1(sK1) <- (~1, 14, 15, 16) [subsumption resolution 435,222] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
440. $false <- (~1, 14, 15, 16) [subsumption resolution 439,142] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
441. 1 | ~14 | ~15 | ~16 [avatar contradiction clause 440]
442. $false [avatar sat refutation 143,148,216,227,338,344,390,411,441] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:194,thDist:-194}
% SZS output end Proof for subset_1__t45_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 33
% Introduced skolems: 15
% 
% >>> Saturation
% Initial clauses: 46
% Generated clauses: 231
% Active clauses: 71
% Passive clauses: 106
% Final active clauses: 65
% Final passive clauses: 32
% Main loop iterations started: 74
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Fw subsumption resolutions: 25
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Forward subsumptions: 46
% Backward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 104
% Unit resulting resolution: 2
% Forward superposition: 2
% Backward superposition: 11
% Self superposition: 3
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 35
% SAT solver unit clauses: 10
% SAT solver binary clauses: 12
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

