i: [1,224,0,2,1,1,0,0] t125_xreal_1
a: [3,360,0,2,1,224]
s: 360
i: [1,223,0,2,1,1,0,0] t125_xreal_1
a: [3,365,0,2,1,223]
s: 365
i: [1,227,0,3,1,1,0,0] t125_xreal_1
a: [3,345,0,3,1,227]
s: 345
i: [1,226,0,4,1,1,0,0] t125_xreal_1
a: [3,350,0,4,1,226]
s: 350
i: [1,225,0,4,1,1,0,0] t125_xreal_1
a: [3,355,0,4,1,225]
s: 355
i: [1,228,0,2,1,0,0,255] fc1_numbers
a: [3,370,0,2,1,228]
s: 370
i: [1,229,0,2,1,0,0,2] fc1_xboole_0
a: [3,375,0,2,1,229]
s: 375
i: [1,236,0,2,1,0,0,3] fc6_ordinal1
a: [3,405,0,2,1,236]
s: 405
i: [1,235,0,2,1,0,0,3] fc6_ordinal1
a: [3,410,0,2,1,235]
s: 410
i: [1,309,0,2,1,0,0,2] rc2_xboole_0
a: [3,422,0,2,1,309]
s: 422
i: [1,310,0,2,1,0,0,4] rc4_ordinal1
a: [3,427,0,2,1,310]
s: 427
i: [1,311,0,2,1,0,0,4] rc1_ordinal1
a: [3,432,0,2,1,311]
s: 432
i: [1,312,0,2,1,0,0,1] rc1_xreal_0
a: [3,437,0,2,1,312]
s: 437
i: [1,313,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,442,0,2,1,313]
s: 442
i: [1,314,0,2,1,0,0,2] rc1_xxreal_0
a: [3,447,0,2,1,314]
s: 447
i: [1,315,0,2,1,0,0,2] rc1_xboole_0
a: [3,452,0,2,1,315]
s: 452
i: [1,319,0,2,1,0,0,255] rc2_ordinal1
a: [3,457,0,2,1,319]
s: 457
i: [1,318,0,2,1,0,0,255] rc2_ordinal1
a: [3,462,0,2,1,318]
s: 462
i: [1,317,0,2,1,0,0,255] rc2_ordinal1
a: [3,467,0,2,1,317]
s: 467
i: [1,316,0,2,1,0,0,255] rc2_ordinal1
a: [3,472,0,2,1,316]
s: 472
i: [1,321,0,2,1,0,0,4] rc5_ordinal1
a: [3,477,0,2,1,321]
s: 477
i: [1,320,0,2,1,0,0,4] rc5_ordinal1
a: [3,482,0,2,1,320]
s: 482
i: [1,323,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,487,0,2,1,323]
s: 487
i: [1,322,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,492,0,2,1,322]
s: 492
i: [1,327,0,2,1,0,0,2] rc2_xreal_0
a: [3,497,0,2,1,327]
s: 497
i: [1,326,0,2,1,0,0,2] rc2_xreal_0
a: [3,502,0,2,1,326]
s: 502
i: [1,325,0,2,1,0,0,2] rc2_xreal_0
a: [3,507,0,2,1,325]
s: 507
i: [1,324,0,2,1,0,0,2] rc2_xreal_0
a: [3,512,0,2,1,324]
s: 512
i: [1,331,0,2,1,0,0,2] rc3_xreal_0
a: [3,517,0,2,1,331]
s: 517
i: [1,330,0,2,1,0,0,2] rc3_xreal_0
a: [3,522,0,2,1,330]
s: 522
i: [1,329,0,2,1,0,0,2] rc3_xreal_0
a: [3,527,0,2,1,329]
s: 527
i: [1,328,0,2,1,0,0,2] rc3_xreal_0
a: [3,532,0,2,1,328]
s: 532
i: [1,333,0,2,1,0,0,2] rc3_xxreal_0
a: [3,537,0,2,1,333]
s: 537
i: [1,332,0,2,1,0,0,2] rc3_xxreal_0
a: [3,542,0,2,1,332]
s: 542
i: [1,335,0,2,1,0,0,2] rc2_xxreal_0
a: [3,547,0,2,1,335]
s: 547
i: [1,334,0,2,1,0,0,2] rc2_xxreal_0
a: [3,552,0,2,1,334]
s: 552
i: [1,339,0,2,1,0,0,2] rc4_xreal_0
a: [3,557,0,2,1,339]
s: 557
i: [1,338,0,2,1,0,0,2] rc4_xreal_0
a: [3,562,0,2,1,338]
s: 562
i: [1,337,0,2,1,0,0,2] rc4_xreal_0
a: [3,567,0,2,1,337]
s: 567
i: [1,336,0,2,1,0,0,2] rc4_xreal_0
a: [3,572,0,2,1,336]
s: 572
i: [1,341,0,2,1,0,0,2] rc4_xxreal_0
a: [3,577,0,2,1,341]
s: 577
i: [1,340,0,2,1,0,0,2] rc4_xxreal_0
a: [3,582,0,2,1,340]
s: 582
i: [1,301,0,3,1,0,0,255] reflexivity_r1_tarski
s: 301
i: [1,230,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,380,0,3,1,230]
s: 380
i: [1,231,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,385,0,3,1,231]
s: 385
r: 375
d: [2,584,0,2,1,2,23,375,385]
a: [3,588,0,2,1,584]
s: 588
i: [1,232,0,3,1,0,0,2] t1_numerals
a: [3,390,0,3,1,232]
s: 390
d: [2,590,0,3,1,2,22,390,385]
a: [3,594,0,3,1,590]
s: 594
i: [1,238,0,4,2,0,0,4] cc6_ordinal1
s: 238
i: [1,239,0,4,2,0,0,4] cc2_xreal_0
s: 239
i: [1,240,0,4,2,0,0,4] cc2_xcmplx_0
s: 240
d: [2,596,1,2,1,1,40,238,427]
a: [3,601,1,2,1,596]
s: 601
i: [1,241,0,4,2,0,0,4] cc2_xxreal_0
s: 241
d: [2,597,1,2,1,1,40,238,477]
a: [3,606,1,2,1,597]
s: 606
i: [1,242,0,4,2,0,0,255] cc1_ordinal1
s: 242
i: [1,243,0,4,2,0,0,255] cc1_ordinal1
s: 243
d: [2,608,1,2,1,1,40,239,427]
a: [3,613,1,2,1,608]
s: 613
i: [1,245,0,4,2,0,0,2] cc3_xreal_0
s: 245
d: [2,694,1,2,1,1,40,245,365]
a: [3,704,1,2,1,694]
s: 704
i: [1,246,0,4,2,0,0,2] cc4_xreal_0
s: 246
i: [1,260,0,4,2,0,0,255] cc9_ordinal1
s: 260
d: [2,695,1,2,1,1,40,245,360]
a: [3,709,1,2,1,695]
s: 709
i: [1,261,0,4,2,0,0,4] cc7_ordinal1
s: 261
d: [2,717,1,2,1,1,40,246,365]
a: [3,727,1,2,1,717]
s: 727
i: [1,262,0,4,2,0,0,4] cc3_ordinal1
s: 262
i: [1,300,0,4,1,0,0,3] existence_m1_subset_1
s: 300
d: [2,718,1,2,1,1,40,246,360]
a: [3,732,1,2,1,718]
s: 732
i: [1,233,0,4,1,0,0,255] dt_k5_numbers
a: [3,395,0,4,1,233]
s: 395
d: [2,609,1,2,1,1,40,239,477]
a: [3,618,1,2,1,609]
s: 618
i: [1,234,0,4,1,0,0,255] dt_k6_numbers
a: [3,400,0,4,1,234]
s: 400
d: [2,812,0,4,1,2,22,395,380]
a: [3,816,0,4,1,812]
s: 816
d: [2,620,1,2,1,1,40,240,427]
a: [3,625,1,2,1,620]
s: 625
d: [2,822,0,4,1,2,22,400,380]
a: [3,826,0,4,1,822]
s: 826
d: [2,621,1,2,1,1,40,240,477]
a: [3,630,1,2,1,621]
s: 630
i: [1,247,0,5,2,0,0,1] fc4_xreal_0
s: 247
i: [1,248,0,5,2,0,0,1] fc4_xreal_0
s: 248
d: [2,632,1,2,1,1,40,241,427]
a: [3,637,1,2,1,632]
s: 637
i: [1,258,0,5,2,0,0,1] dt_k5_xcmplx_0
s: 258
d: [2,633,1,2,1,1,40,241,477]
a: [3,642,1,2,1,633]
s: 642
i: [1,263,0,5,2,0,0,2] t6_boole
s: 263
i: [1,264,0,5,2,0,0,3] cc8_ordinal1
s: 264
d: [2,644,1,2,1,1,40,242,405]
a: [3,652,1,2,1,644]
s: 652
i: [1,265,0,5,2,0,0,255] cc1_xreal_0
s: 265
d: [2,646,1,2,1,1,40,242,432]
a: [3,662,1,2,1,646]
s: 662
i: [1,266,0,5,2,0,0,255] cc1_xcmplx_0
s: 266
i: [1,297,0,5,2,0,0,255] t7_boole
s: 297
d: [2,669,1,2,1,1,40,243,405]
a: [3,677,1,2,1,669]
s: 677
i: [1,307,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,417,0,5,2,307]
s: 417
d: [2,671,1,2,1,1,40,243,432]
a: [3,687,1,2,1,671]
s: 687
d: [2,830,0,5,2,1,22,263,385]
s: 830
i: [1,277,0,6,3,0,0,255] cc2_ordinal1
s: 277
d: [2,855,1,6,2,2,44,355,830]
s: 855
i: [1,284,0,6,3,0,0,2] cc3_xxreal_0
s: 284
d: [2,860,2,6,2,3,40,855,417]
a: [3,864,2,3,1,860]
s: 864
i: [1,286,0,6,3,0,0,2] cc3_xxreal_0
s: 286
i: [1,287,0,6,3,0,0,2] cc5_xxreal_0
s: 287
d: [2,856,1,6,2,2,44,350,830]
s: 856
i: [1,289,0,6,3,0,0,2] cc5_xxreal_0
s: 289
d: [2,912,2,6,2,3,40,856,417]
a: [3,916,2,3,1,912]
s: 916
i: [1,291,0,6,3,0,0,2] cc7_xxreal_0
s: 291
i: [1,292,0,6,3,0,0,2] cc7_xxreal_0
s: 292
d: [2,697,1,2,1,1,40,245,437]
a: [3,715,1,2,1,697]
s: 715
i: [1,293,0,6,2,0,0,255] t1_subset
s: 293
d: [2,720,1,2,1,1,40,246,437]
a: [3,738,1,2,1,720]
s: 738
i: [1,302,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 302
i: [1,244,0,7,3,0,0,4] cc5_ordinal1
s: 244
d: [2,740,1,2,1,1,40,260,588]
a: [3,747,1,2,1,740]
s: 747
i: [1,259,0,7,2,0,0,2] involutiveness_k5_xcmplx_0
s: 259
d: [2,741,1,2,1,1,40,260,452]
a: [3,752,1,2,1,741]
s: 752
i: [1,267,0,7,3,0,0,1] fc28_xreal_0
s: 267
i: [1,268,0,7,3,0,0,1] fc28_xreal_0
s: 268
d: [2,742,1,2,1,1,40,260,572]
a: [3,757,1,2,1,742]
s: 757
i: [1,269,0,7,3,0,0,1] fc30_xreal_0
s: 269
d: [2,743,1,2,1,1,40,260,582]
a: [3,762,1,2,1,743]
s: 762
i: [1,270,0,7,3,0,0,1] fc30_xreal_0
s: 270
i: [1,271,0,7,3,0,0,1] fc7_xcmplx_0
s: 271
d: [2,764,1,2,1,1,40,261,588]
a: [3,771,1,2,1,764]
s: 771
i: [1,272,0,7,3,0,0,1] fc7_xcmplx_0
s: 272
d: [2,765,1,2,1,1,40,261,452]
a: [3,776,1,2,1,765]
s: 776
i: [1,278,0,7,3,0,0,1] fc29_xreal_0
s: 278
i: [1,279,0,7,3,0,0,1] fc29_xreal_0
s: 279
d: [2,766,1,2,1,1,40,261,572]
a: [3,781,1,2,1,766]
s: 781
i: [1,280,0,7,3,0,0,1] fc27_xreal_0
s: 280
d: [2,767,1,2,1,1,40,261,582]
a: [3,786,1,2,1,767]
s: 786
i: [1,281,0,7,3,0,0,1] fc27_xreal_0
s: 281
i: [1,295,0,7,2,0,0,255] t3_subset
s: 295
d: [2,788,1,2,1,1,40,262,588]
a: [3,795,1,2,1,788]
s: 795
i: [1,296,0,7,3,0,0,2] t8_boole
s: 296
d: [2,789,1,2,1,1,40,262,452]
a: [3,800,1,2,1,789]
s: 800
i: [1,237,0,8,3,0,0,255] cc10_ordinal1
s: 237
i: [1,274,0,8,4,0,0,2] cc6_xxreal_0
s: 274
d: [2,1043,2,6,3,1,40,274,727]
a: [3,1058,2,2,1,1043]
s: 1058
r: 1058
i: [1,276,0,8,4,0,0,2] cc4_xxreal_0
s: 276
a: [3,1057,2,2,1,1043]
s: 1057
i: [1,282,0,8,4,0,0,2] cc8_xxreal_0
s: 282
i: [1,294,0,8,3,0,0,255] t2_subset
s: 294
a: [3,1062,2,2,1,1043]
s: 1062
i: [1,299,0,9,3,0,0,255] t5_subset
s: 299
a: [3,1065,2,2,1,1043]
s: 1065
i: [1,298,0,10,3,0,0,255] t4_subset
s: 298
i: [1,308,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 308
d: [2,1044,2,6,3,1,40,274,732]
a: [3,1071,2,2,1,1044]
s: 1071
r: 1071
i: [1,251,0,11,5,0,0,1] t2_real
s: 251
a: [3,1070,2,2,1,1044]
s: 1070
i: [1,252,0,11,5,0,0,1] t1_real
s: 252
i: [1,253,0,11,5,0,0,1] t8_real
s: 253
a: [3,1075,2,2,1,1044]
s: 1075
i: [1,254,0,11,5,0,0,1] t7_real
s: 254
a: [3,1078,2,2,1,1044]
s: 1078
i: [1,256,0,11,5,0,0,1] t4_real
s: 256
i: [1,257,0,11,5,0,0,1] t3_real
s: 257
d: [2,1221,1,10,4,1,40,253,350]
a: [3,1232,1,3,1,1221]
s: 1232
d: [2,1267,2,0,0,2,52,360,1232,248]
e: 1267
d: [2,1268,2,2,1,1,40,1232,248]
d: [2,1270,2,0,0,2,21,1268,360]
e: 1270
r: 1232
i: [1,249,0,13,6,0,0,1] t6_real
s: 249
a: [3,1231,1,3,1,1221]
s: 1231
d: [2,1285,3,0,0,2,52,916,1231,246]
e: 1285
d: [2,1286,2,3,1,1,40,1231,246]
d: [2,1289,2,0,0,2,21,1286,916]
e: 1289
r: 916
i: [1,250,0,13,6,0,0,1] t5_real
s: 250
i: [1,306,0,14,4,0,0,255] existence_m2_subset_1
s: 306
a: [3,1236,1,3,1,1221]
s: 1236
d: [2,1313,3,0,0,3,52,1078,360,1236,268]
e: 1313
d: [2,1314,2,4,2,1,40,1236,268]
d: [2,1317,2,2,1,2,21,1314,360]
d: [2,1318,2,0,0,3,21,1317,1078]
e: 1318
r: 1062
r: 1065
r: 1075
r: 1078
i: [1,255,0,15,5,0,0,1] t124_xreal_1
s: 255
a: [3,1061,2,2,1,1043]
s: 1061
i: [1,303,0,15,5,0,0,255] dt_m2_subset_1
s: 303
i: [1,304,0,15,5,0,0,255] redefinition_m2_subset_1
s: 304
a: [3,1066,2,2,1,1043]
s: 1066
i: [1,305,0,15,5,0,0,255] redefinition_m2_subset_1
s: 305
a: [3,1074,2,2,1,1044]
s: 1074
d: [2,1223,1,10,4,1,40,253,355]
a: [3,1249,1,3,1,1223]
s: 1249
d: [2,1371,2,0,0,2,52,365,1249,248]
e: 1371
d: [2,1372,2,2,1,1,40,1249,248]
d: [2,1374,2,0,0,2,21,1372,365]
e: 1374
r: 1249
a: [3,1248,1,3,1,1223]
s: 1248
d: [2,1376,3,0,0,2,52,864,1248,246]
e: 1376
d: [2,1377,2,3,1,1,40,1248,246]
d: [2,1380,2,0,0,2,21,1377,864]
e: 1380
r: 864
a: [3,1079,2,2,1,1044]
s: 1079
a: [3,1257,1,3,1,1223]
s: 1257
d: [2,1398,3,0,0,3,52,1061,365,1257,270]
e: 1398
d: [2,1399,2,4,2,1,40,1257,270]
d: [2,1403,2,2,1,2,21,1399,365]
d: [2,1404,2,0,0,3,21,1403,1061]
e: 1404
r: 1061
r: 1066
r: 1074
r: 1079
r: 1236
a: [3,1240,1,2,1,1221]
s: 1240
r: 1240
r: 1257
a: [3,1239,1,2,1,1221]
s: 1239
a: [3,1253,1,2,1,1223]
s: 1253
d: [2,831,1,2,1,1,40,264,594]
d: [2,935,2,2,1,1,40,771,241]
d: [2,937,2,2,1,1,40,771,239]
d: [2,1168,3,4,2,1,40,1062,286]
d: [2,1169,3,2,1,2,21,1168,727]
d: [2,1211,1,8,4,1,40,252,345]
d: [2,1323,1,6,3,2,21,1211,365]
d: [2,1324,1,4,2,3,21,1323,360]
d: [2,1413,2,4,2,1,40,1240,287]
d: [2,1416,2,2,1,2,21,1413,588]
d: [2,1420,2,4,2,1,40,1253,284]
d: [2,1422,2,2,1,2,21,1420,588]
f: 1422,1416,1403,1374,1317,1324,1270,1223,1221,1169,937,935,831,717,590,584,232,231,229,223,224,225,226,227
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t125_xreal_1
% SZS output start Proof for xreal_1__t125_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(X1,X0) & ~r1_xxreal_0(k5_xcmplx_0(X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0))))) [input t125_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(X1,X0) & ~r1_xxreal_0(k5_xcmplx_0(X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0))))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
14. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
19. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
40. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
42. ! [X0] : ((v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => (~v2_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0)))) [input fc28_xreal_0]
44. ! [X0] : ((v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => (~v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0)))) [input fc30_xreal_0]
45. ! [X0] : (v1_xreal_0(X0) => (v1_xreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0)))) [input fc4_xreal_0]
67. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
72. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
73. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
88. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
100. ? [X0] : (? [X1] : ((r1_xxreal_0(X1,X0) & ~r1_xxreal_0(k5_xcmplx_0(X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
101. ? [X0] : (? [X1] : (r1_xxreal_0(X1,X0) & ~r1_xxreal_0(k5_xcmplx_0(X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 100]
104. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
106. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
110. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
111. ! [X0] : ((v1_xreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | ~v1_xreal_0(X0)) [ennf transformation 45]
118. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 73]
119. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 118]
120. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 88]
121. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 120]
136. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
139. ! [X0] : ((~v2_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | (~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 42]
140. ! [X0] : ((~v2_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 139]
141. ! [X0] : ((~v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | (~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 44]
142. ! [X0] : ((~v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 141]
157. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 14]
158. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 157]
159. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 19]
160. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 159]
183. ? [X0] : (? [X1] : (r1_xxreal_0(X1,X0) & ~r1_xxreal_0(k5_xcmplx_0(X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (r1_xxreal_0(X1,sK0) & ~r1_xxreal_0(k5_xcmplx_0(X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
184. ? [X1] : (r1_xxreal_0(X1,sK0) & ~r1_xxreal_0(k5_xcmplx_0(X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) & v1_xreal_0(X1)) => (r1_xxreal_0(sK1,sK0) & ~r1_xxreal_0(k5_xcmplx_0(sK1),k6_numbers) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) & v1_xreal_0(sK1)) [choice axiom]
185. (r1_xxreal_0(sK1,sK0) & ~r1_xxreal_0(k5_xcmplx_0(sK1),k6_numbers) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 101,184,183]
223. v1_xreal_0(sK0) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
224. v1_xreal_0(sK1) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
225. ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) [cnf transformation 185] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
226. ~r1_xxreal_0(k5_xcmplx_0(sK1),k6_numbers) [cnf transformation 185] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
227. r1_xxreal_0(sK1,sK0) [cnf transformation 185] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
229. v1_xboole_0(k1_xboole_0) [cnf transformation 40] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
231. k6_numbers = k1_xboole_0 [cnf transformation 67] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
232. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 72] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
239. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 104] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
241. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 106] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
246. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 110] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
248. v1_xreal_0(k5_xcmplx_0(X0)) | ~v1_xreal_0(X0) [cnf transformation 111] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
252. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 119] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
253. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 121] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
264. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 136] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
268. ~v2_xxreal_0(k5_xcmplx_0(X0)) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 140] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
270. ~v3_xxreal_0(k5_xcmplx_0(X0)) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 142] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
284. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 158] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
286. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 158] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
287. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 160] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
343. 1 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
345. r1_xxreal_0(sK1,sK0) <- (1) [avatar component clause 343] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
346. 1 [avatar split clause 227,343]
348. 2 <=> r1_xxreal_0(k5_xcmplx_0(sK1),k6_numbers) [avatar definition]
350. ~r1_xxreal_0(k5_xcmplx_0(sK1),k6_numbers) <- (~2) [avatar component clause 348] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
351. ~2 [avatar split clause 226,348]
353. 3 <=> r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) [avatar definition]
355. ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) <- (~3) [avatar component clause 353] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
356. ~3 [avatar split clause 225,353]
358. 4 <=> v1_xreal_0(sK1) [avatar definition]
360. v1_xreal_0(sK1) <- (4) [avatar component clause 358] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
361. 4 [avatar split clause 224,358]
363. 5 <=> v1_xreal_0(sK0) [avatar definition]
365. v1_xreal_0(sK0) <- (5) [avatar component clause 363] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
366. 5 [avatar split clause 223,363]
373. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
375. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 373] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
376. 7 [avatar split clause 229,373]
383. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
385. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 383] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
386. 9 [avatar split clause 231,383]
388. 10 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
390. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (10) [avatar component clause 388] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
391. 10 [avatar split clause 232,388]
584. v1_xboole_0(k6_numbers) <- (7, 9) [backward demodulation 375,385] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
586. 50 <=> v1_xboole_0(k6_numbers) [avatar definition]
588. v1_xboole_0(k6_numbers) <- (50) [avatar component clause 586] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
589. 50 | ~7 | ~9 [avatar split clause 584,383,373,586]
590. m1_subset_1(k6_numbers,k4_ordinal1) <- (9, 10) [forward demodulation 390,385] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
592. 51 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
594. m1_subset_1(k6_numbers,k4_ordinal1) <- (51) [avatar component clause 592] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
595. 51 | ~9 | ~10 [avatar split clause 590,388,383,592]
717. v1_xxreal_0(sK0) <- (5) [resolution 246,365] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
725. 71 <=> v1_xxreal_0(sK0) [avatar definition]
727. v1_xxreal_0(sK0) <- (71) [avatar component clause 725] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
728. 71 | ~5 [avatar split clause 717,363,725]
769. 78 <=> v7_ordinal1(k6_numbers) [avatar definition]
771. v7_ordinal1(k6_numbers) <- (78) [avatar component clause 769] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
831. v7_ordinal1(k6_numbers) <- (51) [resolution 264,594] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
833. 78 | ~51 [avatar split clause 831,592,769]
935. v1_xxreal_0(k6_numbers) <- (78) [resolution 771,241] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
937. v1_xreal_0(k6_numbers) <- (78) [resolution 771,239] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
941. 101 <=> v1_xxreal_0(k6_numbers) [avatar definition]
944. 101 | ~78 [avatar split clause 935,769,941]
951. 103 <=> v1_xreal_0(k6_numbers) [avatar definition]
954. 103 | ~78 [avatar split clause 937,769,951]
1060. 116 <=> v3_xxreal_0(sK0) [avatar definition]
1062. v3_xxreal_0(sK0) <- (116) [avatar component clause 1060] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1064. 117 <=> v2_xxreal_0(sK0) [avatar definition]
1077. 120 <=> v2_xxreal_0(sK1) [avatar definition]
1168. ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (116) [resolution 1062,286] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1169. ~v2_xxreal_0(sK0) <- (71, 116) [subsumption resolution 1168,727] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1170. ~117 | ~71 | ~116 [avatar split clause 1169,1060,725,1064]
1211. ~v2_xxreal_0(sK1) | v2_xxreal_0(sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(sK1) <- (1) [resolution 252,345] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1221. v3_xxreal_0(k6_numbers) | v2_xxreal_0(k5_xcmplx_0(sK1)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k5_xcmplx_0(sK1)) <- (~2) [resolution 253,350] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1223. v3_xxreal_0(k5_xcmplx_0(sK0)) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k5_xcmplx_0(sK0)) | ~v1_xreal_0(k6_numbers) <- (~3) [resolution 253,355] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1230. 141 <=> v1_xreal_0(k5_xcmplx_0(sK1)) [avatar definition]
1232. ~v1_xreal_0(k5_xcmplx_0(sK1)) <- (~141) [avatar component clause 1230] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1234. 142 <=> v2_xxreal_0(k5_xcmplx_0(sK1)) [avatar definition]
1236. v2_xxreal_0(k5_xcmplx_0(sK1)) <- (142) [avatar component clause 1234] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1238. 143 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1240. v3_xxreal_0(k6_numbers) <- (143) [avatar component clause 1238] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1241. ~141 | ~103 | 142 | 143 | 2 [avatar split clause 1221,348,1238,1234,951,1230]
1247. 145 <=> v1_xreal_0(k5_xcmplx_0(sK0)) [avatar definition]
1249. ~v1_xreal_0(k5_xcmplx_0(sK0)) <- (~145) [avatar component clause 1247] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1251. 146 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1253. v2_xxreal_0(k6_numbers) <- (146) [avatar component clause 1251] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1255. 147 <=> v3_xxreal_0(k5_xcmplx_0(sK0)) [avatar definition]
1257. v3_xxreal_0(k5_xcmplx_0(sK0)) <- (147) [avatar component clause 1255] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1258. ~103 | ~145 | 146 | 147 | 3 [avatar split clause 1223,353,1255,1251,1247,951]
1268. ~v1_xreal_0(sK1) <- (~141) [resolution 1232,248] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1270. $false <- (4, ~141) [subsumption resolution 1268,360] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1271. ~4 | 141 [avatar contradiction clause 1270]
1314. ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) <- (142) [resolution 1236,268] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1317. v2_xxreal_0(sK1) <- (4, 142) [subsumption resolution 1314,360] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1323. ~v2_xxreal_0(sK1) | v2_xxreal_0(sK0) | ~v1_xreal_0(sK1) <- (1, 5) [subsumption resolution 1211,365] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1324. ~v2_xxreal_0(sK1) | v2_xxreal_0(sK0) <- (1, 4, 5) [subsumption resolution 1323,360] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1325. 117 | ~120 | ~1 | ~4 | ~5 [avatar split clause 1324,363,358,343,1077,1064]
1334. 120 | ~4 | ~142 [avatar split clause 1317,1234,358,1077]
1372. ~v1_xreal_0(sK0) <- (~145) [resolution 1249,248] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1374. $false <- (5, ~145) [subsumption resolution 1372,365] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1375. ~5 | 145 [avatar contradiction clause 1374]
1399. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (147) [resolution 1257,270] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1403. v3_xxreal_0(sK0) <- (5, 147) [subsumption resolution 1399,365] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1406. 116 | ~5 | ~147 [avatar split clause 1403,1255,363,1060]
1413. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (143) [resolution 1240,287] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1416. ~v1_xxreal_0(k6_numbers) <- (50, 143) [subsumption resolution 1413,588] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1417. ~101 | ~50 | ~143 [avatar split clause 1416,1238,586,941]
1420. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (146) [resolution 1253,284] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1422. ~v1_xxreal_0(k6_numbers) <- (50, 146) [subsumption resolution 1420,588] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1423. ~101 | ~50 | ~146 [avatar split clause 1422,1251,586,941]
1424. $false [avatar sat refutation 346,351,356,361,366,376,386,391,589,595,728,833,944,954,1170,1241,1258,1271,1325,1334,1375,1406,1417,1423] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:981,thDist:-981}
% SZS output end Proof for xreal_1__t125_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 87
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 119
% Generated clauses: 574
% Active clauses: 154
% Passive clauses: 270
% Final active clauses: 136
% Final passive clauses: 93
% Main loop iterations started: 171
% 
% >>> Simplifying Inferences
% Duplicate literals: 22
% Fw subsumption resolutions: 49
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 56
% 
% >>> Generating Inferences
% Binary resolution: 177
% Unit resulting resolution: 6
% Factoring: 1
% Forward superposition: 13
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 30
% Split components: 73
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 180
% SAT solver unit clauses: 50
% SAT solver binary clauses: 77
% 
% Memory used [KB]: 5756
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% congruence closure: 0.001 s
% resolution: 0.001 s
% other: 0.008 s ( own 0.004 s ) 

