i: [1,84,0,2,1,1,0,0] l6_cfuncdom
a: [3,144,0,2,1,84]
s: 144
i: [1,83,0,2,1,1,0,0] l6_cfuncdom
a: [3,149,0,2,1,83]
s: 149
i: [1,85,0,3,1,1,0,0] l6_cfuncdom
a: [3,139,0,3,1,85]
s: 139
i: [1,88,0,4,1,1,0,0] l6_cfuncdom
a: [3,124,0,4,1,88]
s: 124
i: [1,86,0,4,1,1,0,0] l6_cfuncdom
a: [3,134,0,4,1,86]
s: 134
i: [1,87,0,6,1,1,0,0] l6_cfuncdom
a: [3,129,0,6,1,87]
s: 129
i: [1,89,0,2,1,0,0,2] fc1_xboole_0
a: [3,154,0,2,1,89]
s: 154
i: [1,116,0,2,1,0,0,1] rc2_xboole_0
a: [3,159,0,2,1,116]
s: 159
i: [1,117,0,2,1,0,0,1] rc1_xboole_0
a: [3,164,0,2,1,117]
s: 164
i: [1,103,0,3,1,0,0,255] reflexivity_r1_tarski
s: 103
i: [1,102,0,4,1,0,0,1] existence_m1_subset_1
s: 102
i: [1,112,0,4,1,0,0,255] rc1_relset_1
s: 112
i: [1,113,0,4,1,0,0,255] rc1_relset_1
s: 113
i: [1,90,0,5,2,0,0,2] t6_boole
s: 90
i: [1,95,0,5,2,0,0,1] t7_boole
s: 95
i: [1,114,0,5,1,0,0,255] rc1_relset_1
s: 114
i: [1,91,0,6,2,0,0,1] t1_subset
s: 91
i: [1,106,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 106
i: [1,93,0,7,2,0,0,255] t3_subset
s: 93
i: [1,94,0,7,3,0,0,2] t8_boole
s: 94
i: [1,92,0,8,3,0,0,1] t2_subset
s: 92
d: [2,169,1,5,2,1,40,92,139]
d: [2,172,1,3,1,2,21,169,149]
a: [3,176,1,3,1,172]
s: 176
i: [1,109,0,8,2,0,0,2] cc1_relset_1
s: 109
d: [2,196,1,2,1,1,40,109,129]
a: [3,201,1,2,1,196]
s: 201
i: [1,111,0,8,1,0,0,255] rc1_relset_1
s: 111
i: [1,101,0,9,3,0,0,1] t5_subset
s: 101
d: [2,187,2,3,1,1,40,176,106]
a: [3,193,2,3,1,187]
s: 193
i: [1,110,0,9,2,0,0,255] cc2_relset_1
s: 110
d: [2,218,1,3,1,1,40,110,129]
a: [3,224,1,3,1,218]
s: 224
i: [1,100,0,10,3,0,0,1] t4_subset
s: 100
i: [1,104,0,10,3,0,0,1] cc4_relset_1
s: 104
d: [2,206,1,7,2,1,40,101,129]
a: [3,215,1,4,1,206]
s: 215
i: [1,105,0,10,3,0,0,1] cc3_relset_1
s: 105
d: [2,170,1,11,2,1,40,92,129]
a: [3,185,1,5,1,170]
s: 185
i: [1,119,0,10,2,0,0,1] d1_funct_2
s: 119
i: [1,107,0,11,3,0,0,1] redefinition_k1_relset_1
s: 107
d: [2,227,1,8,2,1,40,100,129]
s: 227
i: [1,108,0,11,3,0,0,2] dt_k1_relset_1
s: 108
d: [2,240,2,12,3,1,40,227,92]
d: [2,241,2,8,2,2,21,240,215]
s: 241
i: [1,115,0,11,2,0,0,1] fc8_relset_1
s: 115
i: [1,120,0,13,3,0,0,1] d1_funct_2
s: 120
d: [2,251,3,8,2,2,40,241,106]
s: 251
i: [1,96,0,18,4,0,0,1] d1_funct_2
s: 96
d: [2,272,1,12,3,1,40,96,129]
d: [2,275,1,8,2,2,21,272,134]
a: [3,279,1,5,1,275]
s: 279
d: [2,287,2,9,3,1,44,107,279]
d: [2,302,2,6,2,2,21,287,224]
d: [2,303,2,4,1,3,21,302,201]
f: 303,275,218,196,172,89,83,84,85,86,87,88
% Refutation found. Thanks to Tanya!
% SZS status Theorem for cfuncdom__l6_cfuncdom
% SZS output start Proof for cfuncdom__l6_cfuncdom
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,X0) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1) & v1_funct_1(X3)) => r2_hidden(X2,k9_xtuple_0(X3)))))) [input l6_cfuncdom]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,X0) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1) & v1_funct_1(X3)) => r2_hidden(X2,k9_xtuple_0(X3)))))) [negated conjecture 1]
4. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [input cc1_relset_1]
5. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [input cc2_relset_1]
8. ! [X0] : ! [X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ((k1_xboole_0 = X1 => (v1_funct_2(X2,X0,X1) <=> k1_xboole_0 = X2)) & (k1_xboole_0 != X1 => (v1_funct_2(X2,X0,X1) <=> k1_relset_1(X0,X2) = X0)))) [input d1_funct_2]
16. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
21. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => k1_relset_1(X0,X1) = k9_xtuple_0(X1)) [input redefinition_k1_relset_1]
24. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
32. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 24]
36. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ((k1_xboole_0 = X1 => (v1_funct_2(X2,X0,X1) <=> k1_xboole_0 = X2)) & (k1_xboole_0 != X1 => (v1_funct_2(X2,X0,X1) <=> k1_relset_1(X0,X2) = X0)))) [flattening 8]
40. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [flattening 4]
41. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [flattening 5]
43. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v4_relat_1(X2,X0)) [pure predicate removal 41]
45. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,X0) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1)) => r2_hidden(X2,k9_xtuple_0(X3)))))) [pure predicate removal 2]
46. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & (m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1))) & m1_subset_1(X2,X0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [ennf transformation 45]
47. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1)) & m1_subset_1(X2,X0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [flattening 46]
50. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 32]
51. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 50]
55. ! [X0,X1,X2] : ((((v1_funct_2(X2,X0,X1) <=> k1_xboole_0 = X2) | k1_xboole_0 != X1) & ((v1_funct_2(X2,X0,X1) <=> k1_relset_1(X0,X2) = X0) | k1_xboole_0 = X1)) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 36]
62. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 21]
63. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 62]
66. ! [X0,X1,X2] : (v1_relat_1(X2) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 40]
67. ! [X0,X1,X2] : (v4_relat_1(X2,X0) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 43]
69. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1)) & m1_subset_1(X2,X0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,X1))) & v1_funct_2(X3,sK0,X1)) & m1_subset_1(X2,sK0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
70. ? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,X1))) & v1_funct_2(X3,sK0,X1)) & m1_subset_1(X2,sK0)) & ~v1_xboole_0(X1)) => (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(X3,sK0,sK1)) & m1_subset_1(X2,sK0)) & ~v1_xboole_0(sK1)) [choice axiom]
71. ? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(X3,sK0,sK1)) & m1_subset_1(X2,sK0)) => (? [X3] : (~r2_hidden(sK2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(X3,sK0,sK1)) & m1_subset_1(sK2,sK0)) [choice axiom]
72. ? [X3] : (~r2_hidden(sK2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(X3,sK0,sK1)) => (~r2_hidden(sK2,k9_xtuple_0(sK3)) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(sK3,sK0,sK1)) [choice axiom]
73. (((~r2_hidden(sK2,k9_xtuple_0(sK3)) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(sK3,sK0,sK1)) & m1_subset_1(sK2,sK0)) & ~v1_xboole_0(sK1)) & ~v1_xboole_0(sK0) [skolemisation 47,72,71,70,69]
74. ! [X0,X1,X2] : (((((v1_funct_2(X2,X0,X1) | k1_xboole_0 != X2) & (k1_xboole_0 = X2 | ~v1_funct_2(X2,X0,X1))) | k1_xboole_0 != X1) & (((v1_funct_2(X2,X0,X1) | k1_relset_1(X0,X2) != X0) & (k1_relset_1(X0,X2) = X0 | ~v1_funct_2(X2,X0,X1))) | k1_xboole_0 = X1)) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [nnf transformation 55]
83. ~v1_xboole_0(sK0) [cnf transformation 73] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
84. ~v1_xboole_0(sK1) [cnf transformation 73] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
85. m1_subset_1(sK2,sK0) [cnf transformation 73] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
86. v1_funct_2(sK3,sK0,sK1) [cnf transformation 73] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
87. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [cnf transformation 73] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
88. ~r2_hidden(sK2,k9_xtuple_0(sK3)) [cnf transformation 73] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
89. v1_xboole_0(k1_xboole_0) [cnf transformation 16] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
92. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 51] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
96. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~v1_funct_2(X2,X0,X1) | k1_xboole_0 = X1 | k1_relset_1(X0,X2) = X0 [cnf transformation 74] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
107. k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 63] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
109. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v1_relat_1(X2) [cnf transformation 66] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
110. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v4_relat_1(X2,X0) [cnf transformation 67] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
122. 1 <=> r2_hidden(sK2,k9_xtuple_0(sK3)) [avatar definition]
125. ~1 [avatar split clause 88,122]
127. 2 <=> m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [avatar definition]
129. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) <- (2) [avatar component clause 127] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
130. 2 [avatar split clause 87,127]
132. 3 <=> v1_funct_2(sK3,sK0,sK1) [avatar definition]
134. v1_funct_2(sK3,sK0,sK1) <- (3) [avatar component clause 132] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
135. 3 [avatar split clause 86,132]
137. 4 <=> m1_subset_1(sK2,sK0) [avatar definition]
139. m1_subset_1(sK2,sK0) <- (4) [avatar component clause 137] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
140. 4 [avatar split clause 85,137]
142. 5 <=> v1_xboole_0(sK1) [avatar definition]
145. ~5 [avatar split clause 84,142]
147. 6 <=> v1_xboole_0(sK0) [avatar definition]
149. ~v1_xboole_0(sK0) <- (~6) [avatar component clause 147] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
150. ~6 [avatar split clause 83,147]
152. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
155. 7 [avatar split clause 89,152]
169. v1_xboole_0(sK0) | r2_hidden(sK2,sK0) <- (4) [resolution 92,139] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
172. r2_hidden(sK2,sK0) <- (4, ~6) [subsumption resolution 169,149] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
174. 10 <=> r2_hidden(sK2,sK0) [avatar definition]
177. 10 | ~4 | 6 [avatar split clause 172,147,137,174]
196. v1_relat_1(sK3) <- (2) [resolution 109,129] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
199. 14 <=> v1_relat_1(sK3) [avatar definition]
201. v1_relat_1(sK3) <- (14) [avatar component clause 199] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
202. 14 | ~2 [avatar split clause 196,127,199]
218. v4_relat_1(sK3,sK0) <- (2) [resolution 110,129] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
222. 17 <=> v4_relat_1(sK3,sK0) [avatar definition]
224. v4_relat_1(sK3,sK0) <- (17) [avatar component clause 222] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
225. 17 | ~2 [avatar split clause 218,127,222]
272. ~v1_funct_2(sK3,sK0,sK1) | k1_xboole_0 = sK1 | sK0 = k1_relset_1(sK0,sK3) <- (2) [resolution 96,129] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
275. k1_xboole_0 = sK1 | sK0 = k1_relset_1(sK0,sK3) <- (2, 3) [subsumption resolution 272,134] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
277. 19 <=> sK0 = k1_relset_1(sK0,sK3) [avatar definition]
279. sK0 = k1_relset_1(sK0,sK3) <- (19) [avatar component clause 277] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
281. 20 <=> k1_xboole_0 = sK1 [avatar definition]
284. 19 | 20 | ~2 | ~3 [avatar split clause 275,132,127,281,277]
287. sK0 = k9_xtuple_0(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (19) [superposition 107,279] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
291. 21 <=> sK0 = k9_xtuple_0(sK3) [avatar definition]
302. sK0 = k9_xtuple_0(sK3) | ~v1_relat_1(sK3) <- (17, 19) [subsumption resolution 287,224] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
303. sK0 = k9_xtuple_0(sK3) <- (14, 17, 19) [subsumption resolution 302,201] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
304. 21 | ~14 | ~17 | ~19 [avatar split clause 303,277,222,199,291]
305. k1_xboole_0 != sK1 | v1_xboole_0(sK1) | ~v1_xboole_0(k1_xboole_0) [theory tautology sat conflict] {a:0,w:7,wCS:2800,sine:255,thAx:0,allAx:1,thDist:-1}
306. sK0 != k9_xtuple_0(sK3) | r2_hidden(sK2,k9_xtuple_0(sK3)) | ~r2_hidden(sK2,sK0) [theory tautology sat conflict] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
307. $false [avatar sat refutation 125,130,135,140,145,150,155,177,202,225,284,304,305,306] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:159,thDist:-159}
% SZS output end Proof for cfuncdom__l6_cfuncdom
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 35
% Generated clauses: 129
% Active clauses: 44
% Passive clauses: 87
% Final active clauses: 44
% Final passive clauses: 43
% Main loop iterations started: 44
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 58
% Forward superposition: 4
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 9
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

