i: [1,61,0,2,1,1,0,0] t88_funct_4
a: [3,99,0,2,1,61]
s: 99
i: [1,60,0,2,1,1,0,0] t88_funct_4
a: [3,104,0,2,1,60]
s: 104
i: [1,59,0,2,1,1,0,0] t88_funct_4
a: [3,109,0,2,1,59]
s: 109
i: [1,58,0,2,1,1,0,0] t88_funct_4
a: [3,114,0,2,1,58]
s: 114
i: [1,85,0,2,1,0,0,1] rc1_funct_1
a: [3,126,0,2,1,85]
s: 126
i: [1,84,0,2,1,0,0,1] rc1_funct_1
a: [3,131,0,2,1,84]
s: 131
i: [1,62,0,9,1,1,0,0] t88_funct_4
a: [3,94,0,9,1,62]
s: 94
i: [1,63,0,11,1,1,0,0] t88_funct_4
a: [3,89,0,11,1,63]
s: 89
i: [1,73,0,3,1,0,0,1] reflexivity_r1_tarski
s: 73
i: [1,72,0,4,1,0,0,255] existence_m1_subset_1
s: 72
i: [1,74,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 74
i: [1,76,0,6,2,0,0,1] fc1_relat_1
s: 76
i: [1,77,0,6,2,0,0,1] dt_k5_relat_1
s: 77
i: [1,70,0,7,2,0,0,255] t3_subset
s: 70
i: [1,71,0,7,2,0,0,255] t3_subset
s: 71
i: [1,75,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 75
d: [2,134,1,5,1,0,40,70,72]
s: 134
i: [1,64,0,8,2,0,0,1] rd4_relat_1
s: 64
d: [2,136,1,6,2,0,44,76,75]
s: 136
i: [1,65,0,8,3,0,0,255] cc2_relat_1
s: 65
i: [1,82,0,8,3,0,0,1] fc8_funct_1
s: 82
d: [2,137,1,6,2,0,44,76,75]
s: 137
i: [1,83,0,8,3,0,0,1] fc8_funct_1
s: 83
d: [2,143,1,6,2,0,40,65,72]
s: 143
i: [1,79,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,121,0,9,3,79]
s: 121
i: [1,66,0,10,4,0,0,255] cc3_funct_1
s: 66
d: [2,144,1,7,3,0,40,65,71]
s: 144
i: [1,69,0,11,2,0,0,1] t61_relat_1
s: 69
d: [2,149,2,10,2,1,40,144,94]
a: [3,159,2,5,1,149]
s: 159
i: [1,78,0,11,2,0,0,1] rd5_relat_1
s: 78
i: [1,80,0,12,5,0,0,1] dt_k1_funct_4
s: 80
d: [2,170,3,2,1,1,40,159,76]
a: [3,179,3,2,1,170]
s: 179
i: [1,81,0,12,5,0,0,1] dt_k1_funct_4
s: 81
d: [2,169,3,3,1,1,40,159,136]
a: [3,174,3,3,1,169]
s: 174
i: [1,67,0,18,6,0,0,1] t19_funct_4
s: 67
i: [1,68,0,21,5,0,0,1] t71_funct_4
s: 68
d: [2,147,1,8,3,0,40,66,72]
s: 147
d: [2,148,1,9,4,0,40,66,71]
s: 148
d: [2,165,2,9,3,0,44,136,69]
s: 165
d: [2,161,1,10,2,0,44,69,74]
s: 161
d: [2,164,1,10,2,0,44,74,69]
s: 164
d: [2,168,1,10,3,0,44,76,69]
s: 168
d: [2,162,1,11,2,0,44,69,75]
s: 162
d: [2,246,3,7,2,1,44,159,162]
d: [2,261,3,5,1,2,21,246,104]
a: [3,265,3,5,1,261]
s: 265
d: [2,163,1,11,2,0,44,69,75]
s: 163
d: [2,166,1,11,2,0,44,75,69]
s: 166
d: [2,245,2,11,2,1,44,94,162]
d: [2,255,2,9,1,2,21,245,114]
a: [3,259,2,9,1,255]
s: 259
d: [2,167,1,11,2,0,44,75,69]
s: 167
d: [2,247,2,11,2,1,44,94,162]
d: [2,267,2,9,1,2,21,247,104]
a: [3,271,2,9,1,267]
s: 271
d: [2,196,1,25,7,1,44,68,121]
d: [2,201,1,21,5,1,18,196]
d: [2,206,1,17,4,1,21,201,77]
d: [2,207,1,13,3,1,21,206,83]
s: 207
d: [2,197,1,25,7,1,44,121,68]
d: [2,200,1,21,5,1,18,197]
d: [2,208,1,17,4,1,21,200,77]
d: [2,209,1,13,3,1,21,208,83]
s: 209
d: [2,277,3,11,2,1,44,259,162]
d: [2,278,3,9,1,2,21,277,104]
a: [3,282,3,9,1,278]
s: 282
d: [2,198,1,30,9,0,44,81,68]
d: [2,210,1,26,8,0,21,198,83]
d: [2,211,1,22,7,0,21,210,83]
d: [2,212,1,18,6,0,21,211,77]
d: [2,213,1,14,5,0,21,212,77]
s: 213
d: [2,311,4,27,5,1,40,282,67]
a: [3,317,4,4,1,311]
s: 317
d: [2,345,5,0,0,2,52,114,317,77]
e: 345
d: [2,346,5,2,1,1,40,317,77]
d: [2,348,5,0,0,2,21,346,114]
e: 348
r: 317
d: [2,199,1,30,9,0,44,80,68]
d: [2,214,1,26,8,0,21,199,77]
d: [2,215,1,22,7,0,21,214,77]
d: [2,216,1,18,6,0,21,215,83]
d: [2,217,1,14,5,0,21,216,83]
s: 217
d: [2,194,1,23,6,0,44,68,64]
d: [2,202,1,21,5,0,18,194]
s: 202
a: [3,316,4,4,1,311]
s: 316
d: [2,192,1,23,6,0,44,68,64]
d: [2,203,1,21,5,0,18,192]
s: 203
a: [3,321,4,4,1,311]
s: 321
d: [2,402,5,0,0,3,52,114,109,321,83]
e: 402
d: [2,403,5,4,2,1,40,321,83]
d: [2,405,5,2,1,2,21,403,109]
d: [2,406,5,0,0,3,21,405,114]
e: 406
r: 321
d: [2,193,1,29,6,0,44,68,78]
d: [2,204,1,25,5,0,21,193,77]
s: 204
d: [2,195,1,29,6,0,44,68,78]
d: [2,205,1,25,5,0,21,195,77]
s: 205
a: [3,320,4,4,1,311]
s: 320
d: [2,224,2,18,3,0,44,161,161]
d: [2,231,2,13,2,0,21,224,77]
s: 231
a: [3,325,4,4,1,311]
s: 325
d: [2,521,5,0,0,2,52,104,325,77]
e: 521
d: [2,522,5,2,1,1,40,325,77]
d: [2,524,5,0,0,2,21,522,104]
e: 524
r: 325
d: [2,228,2,19,3,0,44,69,161]
d: [2,234,2,14,2,0,21,228,77]
s: 234
d: [2,240,2,19,3,0,44,162,161]
d: [2,254,2,14,2,0,21,240,77]
s: 254
a: [3,324,4,4,1,311]
s: 324
d: [2,290,2,18,4,1,44,207,64]
d: [2,303,2,14,3,1,21,290,80]
s: 303
a: [3,329,4,4,1,311]
s: 329
d: [2,662,5,0,0,3,52,104,99,329,83]
e: 662
d: [2,663,5,4,2,1,40,329,83]
d: [2,665,5,2,1,2,21,663,99]
d: [2,666,5,0,0,3,21,665,104]
e: 666
r: 329
d: [2,297,2,18,5,1,44,168,207]
d: [2,308,2,14,4,1,21,297,80]
s: 308
d: [2,299,2,18,4,1,44,64,207]
d: [2,310,2,14,3,1,21,299,80]
s: 310
a: [3,328,4,4,1,311]
s: 328
d: [2,361,2,20,6,0,44,202,64]
d: [2,372,2,14,3,0,18,361]
s: 372
a: [3,333,4,11,1,311]
s: 333
d: [2,741,5,21,5,1,44,333,205]
d: [2,751,5,10,4,2,21,741,89]
d: [2,752,5,6,3,3,21,751,328]
d: [2,753,5,4,2,4,21,752,109]
d: [2,754,5,2,1,5,21,753,114]
d: [2,755,5,0,0,6,21,754,104]
e: 755
d: [2,744,5,21,5,1,44,205,333]
d: [2,765,5,10,4,2,21,744,89]
d: [2,766,5,6,3,3,21,765,328]
d: [2,767,5,4,2,4,21,766,109]
d: [2,768,5,2,1,5,21,767,114]
d: [2,769,5,0,0,6,21,768,104]
e: 769
d: [2,286,3,11,2,1,44,271,162]
d: [2,287,3,9,1,2,21,286,114]
f: 769,666,524,406,348,311,287,267,58,59,60,61,62,63
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t88_funct_4
% SZS output start Proof for funct_4__t88_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : (r1_tarski(k3_xboole_0(X2,k9_xtuple_0(X0)),k3_xboole_0(X2,k9_xtuple_0(X1))) => k5_relat_1(X1,X2) = k5_relat_1(k1_funct_4(X0,k5_relat_1(X1,X2)),X2)))) [input t88_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : (r1_tarski(k3_xboole_0(X2,k9_xtuple_0(X0)),k3_xboole_0(X2,k9_xtuple_0(X1))) => k5_relat_1(X1,X2) = k5_relat_1(k1_funct_4(X0,k5_relat_1(X1,X2)),X2)))) [negated conjecture 1]
5. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
9. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k5_relat_1(X0,X1))) [input dt_k5_relat_1]
14. ! [X0,X1] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1)))) [input fc8_funct_1]
19. ! [X0,X1] : (v1_relat_1(X0) => k5_relat_1(X0,X1) = k5_relat_1(k5_relat_1(X0,X1),X1)) [input rd5_relat_1]
21. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k1_funct_4(X0,X1) = X1))) [input t19_funct_4]
23. ! [X0] : ! [X1] : (v1_relat_1(X1) => k9_xtuple_0(k5_relat_1(X1,X0)) = k3_xboole_0(k9_xtuple_0(X1),X0)) [input t61_relat_1]
24. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : k5_relat_1(k1_funct_4(X0,X1),X2) = k1_funct_4(k5_relat_1(X0,X2),k5_relat_1(X1,X2)))) [input t71_funct_4]
25. ! [X0,X1] : (v1_relat_1(X1) => k9_xtuple_0(k5_relat_1(X1,X0)) = k3_xboole_0(k9_xtuple_0(X1),X0)) [flattening 23]
29. ? [X0] : (? [X1] : (? [X2] : (k5_relat_1(X1,X2) != k5_relat_1(k1_funct_4(X0,k5_relat_1(X1,X2)),X2) & r1_tarski(k3_xboole_0(X2,k9_xtuple_0(X0)),k3_xboole_0(X2,k9_xtuple_0(X1)))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
30. ? [X0] : (? [X1] : (? [X2] : (k5_relat_1(X1,X2) != k5_relat_1(k1_funct_4(X0,k5_relat_1(X1,X2)),X2) & r1_tarski(k3_xboole_0(X2,k9_xtuple_0(X0)),k3_xboole_0(X2,k9_xtuple_0(X1)))) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 29]
35. ! [X0] : (! [X1] : ((k1_funct_4(X0,X1) = X1 | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 21]
36. ! [X0] : (! [X1] : (k1_funct_4(X0,X1) = X1 | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 35]
37. ! [X0] : (! [X1] : (! [X2] : k5_relat_1(k1_funct_4(X0,X1),X2) = k1_funct_4(k5_relat_1(X0,X2),k5_relat_1(X1,X2)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 24]
38. ! [X0] : (! [X1] : (! [X2] : k5_relat_1(k1_funct_4(X0,X1),X2) = k1_funct_4(k5_relat_1(X0,X2),k5_relat_1(X1,X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 37]
39. ! [X0,X1] : (k9_xtuple_0(k5_relat_1(X1,X0)) = k3_xboole_0(k9_xtuple_0(X1),X0) | ~v1_relat_1(X1)) [ennf transformation 25]
41. ! [X0,X1] : (v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 9]
42. ! [X0,X1] : (k5_relat_1(X0,X1) = k5_relat_1(k5_relat_1(X0,X1),X1) | ~v1_relat_1(X0)) [ennf transformation 19]
47. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 14]
48. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 47]
49. ? [X0] : (? [X1] : (? [X2] : (k5_relat_1(X1,X2) != k5_relat_1(k1_funct_4(X0,k5_relat_1(X1,X2)),X2) & r1_tarski(k3_xboole_0(X2,k9_xtuple_0(X0)),k3_xboole_0(X2,k9_xtuple_0(X1)))) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (k5_relat_1(X1,X2) != k5_relat_1(k1_funct_4(sK0,k5_relat_1(X1,X2)),X2) & r1_tarski(k3_xboole_0(X2,k9_xtuple_0(sK0)),k3_xboole_0(X2,k9_xtuple_0(X1)))) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
50. ? [X1] : (? [X2] : (k5_relat_1(X1,X2) != k5_relat_1(k1_funct_4(sK0,k5_relat_1(X1,X2)),X2) & r1_tarski(k3_xboole_0(X2,k9_xtuple_0(sK0)),k3_xboole_0(X2,k9_xtuple_0(X1)))) & v1_funct_1(X1) & v1_relat_1(X1)) => (? [X2] : (k5_relat_1(sK1,X2) != k5_relat_1(k1_funct_4(sK0,k5_relat_1(sK1,X2)),X2) & r1_tarski(k3_xboole_0(X2,k9_xtuple_0(sK0)),k3_xboole_0(X2,k9_xtuple_0(sK1)))) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
51. ? [X2] : (k5_relat_1(sK1,X2) != k5_relat_1(k1_funct_4(sK0,k5_relat_1(sK1,X2)),X2) & r1_tarski(k3_xboole_0(X2,k9_xtuple_0(sK0)),k3_xboole_0(X2,k9_xtuple_0(sK1)))) => (k5_relat_1(sK1,sK2) != k5_relat_1(k1_funct_4(sK0,k5_relat_1(sK1,sK2)),sK2) & r1_tarski(k3_xboole_0(sK2,k9_xtuple_0(sK0)),k3_xboole_0(sK2,k9_xtuple_0(sK1)))) [choice axiom]
52. ((k5_relat_1(sK1,sK2) != k5_relat_1(k1_funct_4(sK0,k5_relat_1(sK1,sK2)),sK2) & r1_tarski(k3_xboole_0(sK2,k9_xtuple_0(sK0)),k3_xboole_0(sK2,k9_xtuple_0(sK1)))) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 30,51,50,49]
58. v1_relat_1(sK0) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
59. v1_funct_1(sK0) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
60. v1_relat_1(sK1) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
61. v1_funct_1(sK1) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
62. r1_tarski(k3_xboole_0(sK2,k9_xtuple_0(sK0)),k3_xboole_0(sK2,k9_xtuple_0(sK1))) [cnf transformation 52] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
63. k5_relat_1(sK1,sK2) != k5_relat_1(k1_funct_4(sK0,k5_relat_1(sK1,sK2)),sK2) [cnf transformation 52] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
67. ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | k1_funct_4(X0,X1) = X1 | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. k5_relat_1(k1_funct_4(X0,X1),X2) = k1_funct_4(k5_relat_1(X0,X2),k5_relat_1(X1,X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 38] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. k9_xtuple_0(k5_relat_1(X1,X0)) = k3_xboole_0(k9_xtuple_0(X1),X0) | ~v1_relat_1(X1) [cnf transformation 39] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
75. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 5] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
77. v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 41] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
78. k5_relat_1(X0,X1) = k5_relat_1(k5_relat_1(X0,X1),X1) | ~v1_relat_1(X0) [cnf transformation 42] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
83. v1_funct_1(k5_relat_1(X0,X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 48] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. 1 <=> k5_relat_1(sK1,sK2) = k5_relat_1(k1_funct_4(sK0,k5_relat_1(sK1,sK2)),sK2) [avatar definition]
89. k5_relat_1(sK1,sK2) != k5_relat_1(k1_funct_4(sK0,k5_relat_1(sK1,sK2)),sK2) <- (~1) [avatar component clause 87] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
90. ~1 [avatar split clause 63,87]
92. 2 <=> r1_tarski(k3_xboole_0(sK2,k9_xtuple_0(sK0)),k3_xboole_0(sK2,k9_xtuple_0(sK1))) [avatar definition]
94. r1_tarski(k3_xboole_0(sK2,k9_xtuple_0(sK0)),k3_xboole_0(sK2,k9_xtuple_0(sK1))) <- (2) [avatar component clause 92] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
95. 2 [avatar split clause 62,92]
97. 3 <=> v1_funct_1(sK1) [avatar definition]
99. v1_funct_1(sK1) <- (3) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
100. 3 [avatar split clause 61,97]
102. 4 <=> v1_relat_1(sK1) [avatar definition]
104. v1_relat_1(sK1) <- (4) [avatar component clause 102] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
105. 4 [avatar split clause 60,102]
107. 5 <=> v1_funct_1(sK0) [avatar definition]
109. v1_funct_1(sK0) <- (5) [avatar component clause 107] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
110. 5 [avatar split clause 59,107]
112. 6 <=> v1_relat_1(sK0) [avatar definition]
114. v1_relat_1(sK0) <- (6) [avatar component clause 112] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
115. 6 [avatar split clause 58,112]
162. k3_xboole_0(X4,k9_xtuple_0(X3)) = k9_xtuple_0(k5_relat_1(X3,X4)) | ~v1_relat_1(X3) [superposition 69,75] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
195. k5_relat_1(k1_funct_4(X4,k5_relat_1(X2,X3)),X3) = k1_funct_4(k5_relat_1(X4,X3),k5_relat_1(X2,X3)) | ~v1_funct_1(k5_relat_1(X2,X3)) | ~v1_relat_1(k5_relat_1(X2,X3)) | ~v1_funct_1(X4) | ~v1_relat_1(X4) | ~v1_relat_1(X2) [superposition 68,78] {a:1,w:29,wCS:11600,sine:1,thAx:0,allAx:2,thDist:-2}
205. k5_relat_1(k1_funct_4(X4,k5_relat_1(X2,X3)),X3) = k1_funct_4(k5_relat_1(X4,X3),k5_relat_1(X2,X3)) | ~v1_funct_1(k5_relat_1(X2,X3)) | ~v1_funct_1(X4) | ~v1_relat_1(X4) | ~v1_relat_1(X2) [subsumption resolution 195,77] {a:1,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
247. r1_tarski(k3_xboole_0(sK2,k9_xtuple_0(sK0)),k9_xtuple_0(k5_relat_1(sK1,sK2))) | ~v1_relat_1(sK1) <- (2) [superposition 94,162] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
267. r1_tarski(k3_xboole_0(sK2,k9_xtuple_0(sK0)),k9_xtuple_0(k5_relat_1(sK1,sK2))) <- (2, 4) [subsumption resolution 247,104] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
269. 17 <=> r1_tarski(k3_xboole_0(sK2,k9_xtuple_0(sK0)),k9_xtuple_0(k5_relat_1(sK1,sK2))) [avatar definition]
271. r1_tarski(k3_xboole_0(sK2,k9_xtuple_0(sK0)),k9_xtuple_0(k5_relat_1(sK1,sK2))) <- (17) [avatar component clause 269] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
272. 17 | ~2 | ~4 [avatar split clause 267,102,92,269]
280. 18 <=> r1_tarski(k9_xtuple_0(k5_relat_1(sK0,sK2)),k9_xtuple_0(k5_relat_1(sK1,sK2))) [avatar definition]
282. r1_tarski(k9_xtuple_0(k5_relat_1(sK0,sK2)),k9_xtuple_0(k5_relat_1(sK1,sK2))) <- (18) [avatar component clause 280] {a:3,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
286. r1_tarski(k9_xtuple_0(k5_relat_1(sK0,sK2)),k9_xtuple_0(k5_relat_1(sK1,sK2))) | ~v1_relat_1(sK0) <- (17) [superposition 271,162] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
287. r1_tarski(k9_xtuple_0(k5_relat_1(sK0,sK2)),k9_xtuple_0(k5_relat_1(sK1,sK2))) <- (6, 17) [subsumption resolution 286,114] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
288. 18 | ~6 | ~17 [avatar split clause 287,269,112,280]
311. k5_relat_1(sK1,sK2) = k1_funct_4(k5_relat_1(sK0,sK2),k5_relat_1(sK1,sK2)) | ~v1_funct_1(k5_relat_1(sK1,sK2)) | ~v1_relat_1(k5_relat_1(sK1,sK2)) | ~v1_funct_1(k5_relat_1(sK0,sK2)) | ~v1_relat_1(k5_relat_1(sK0,sK2)) <- (18) [resolution 282,67] {a:4,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
315. 19 <=> v1_relat_1(k5_relat_1(sK0,sK2)) [avatar definition]
317. ~v1_relat_1(k5_relat_1(sK0,sK2)) <- (~19) [avatar component clause 315] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
319. 20 <=> v1_funct_1(k5_relat_1(sK0,sK2)) [avatar definition]
321. ~v1_funct_1(k5_relat_1(sK0,sK2)) <- (~20) [avatar component clause 319] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
323. 21 <=> v1_relat_1(k5_relat_1(sK1,sK2)) [avatar definition]
325. ~v1_relat_1(k5_relat_1(sK1,sK2)) <- (~21) [avatar component clause 323] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
327. 22 <=> v1_funct_1(k5_relat_1(sK1,sK2)) [avatar definition]
328. v1_funct_1(k5_relat_1(sK1,sK2)) <- (22) [avatar component clause 327] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
329. ~v1_funct_1(k5_relat_1(sK1,sK2)) <- (~22) [avatar component clause 327] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
331. 23 <=> k5_relat_1(sK1,sK2) = k1_funct_4(k5_relat_1(sK0,sK2),k5_relat_1(sK1,sK2)) [avatar definition]
333. k5_relat_1(sK1,sK2) = k1_funct_4(k5_relat_1(sK0,sK2),k5_relat_1(sK1,sK2)) <- (23) [avatar component clause 331] {a:4,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
334. ~19 | ~20 | ~21 | ~22 | 23 | ~18 [avatar split clause 311,280,331,327,323,319,315]
346. ~v1_relat_1(sK0) <- (~19) [resolution 317,77] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
348. $false <- (6, ~19) [subsumption resolution 346,114] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
349. ~6 | 19 [avatar contradiction clause 348]
403. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~20) [resolution 321,83] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
405. ~v1_relat_1(sK0) <- (5, ~20) [subsumption resolution 403,109] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
406. $false <- (5, 6, ~20) [subsumption resolution 405,114] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
407. ~5 | ~6 | 20 [avatar contradiction clause 406]
522. ~v1_relat_1(sK1) <- (~21) [resolution 325,77] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
524. $false <- (4, ~21) [subsumption resolution 522,104] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
525. ~4 | 21 [avatar contradiction clause 524]
663. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~22) [resolution 329,83] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
665. ~v1_relat_1(sK1) <- (3, ~22) [subsumption resolution 663,99] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
666. $false <- (3, 4, ~22) [subsumption resolution 665,104] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
667. ~3 | ~4 | 22 [avatar contradiction clause 666]
744. k5_relat_1(sK1,sK2) = k5_relat_1(k1_funct_4(sK0,k5_relat_1(sK1,sK2)),sK2) | ~v1_funct_1(k5_relat_1(sK1,sK2)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_relat_1(sK1) <- (23) [superposition 205,333] {a:5,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
765. ~v1_funct_1(k5_relat_1(sK1,sK2)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_relat_1(sK1) <- (~1, 23) [subsumption resolution 744,89] {a:5,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
766. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_relat_1(sK1) <- (~1, 22, 23) [subsumption resolution 765,328] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
767. ~v1_relat_1(sK0) | ~v1_relat_1(sK1) <- (~1, 5, 22, 23) [subsumption resolution 766,109] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
768. ~v1_relat_1(sK1) <- (~1, 5, 6, 22, 23) [subsumption resolution 767,114] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
769. $false <- (~1, 4, 5, 6, 22, 23) [subsumption resolution 768,104] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
770. 1 | ~4 | ~5 | ~6 | ~22 | ~23 [avatar contradiction clause 769]
779. $false [avatar sat refutation 90,95,100,105,110,115,272,288,334,349,407,525,667,770] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:259,thDist:-259}
% SZS output end Proof for funct_4__t88_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 28
% Generated clauses: 628
% Active clauses: 66
% Passive clauses: 209
% Final active clauses: 62
% Final passive clauses: 135
% Main loop iterations started: 74
% 
% >>> Simplifying Inferences
% Duplicate literals: 175
% Fw subsumption resolutions: 194
% Fw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 24
% Forward subsumptions: 77
% 
% >>> Generating Inferences
% Binary resolution: 28
% Unit resulting resolution: 4
% Forward superposition: 106
% Backward superposition: 138
% Self superposition: 5
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 9
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 30
% SAT solver unit clauses: 8
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9850
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward superposition index maintenance: 0.001 s
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.001 s
% superposition: 0.001 s
% other: 0.006 s ( own 0.001 s ) 

