i: [1,32,0,3,1,1,0,0] t6_xboole_0
a: [3,50,0,3,1,32]
s: 50
i: [1,33,0,6,2,1,0,0] t6_xboole_0
s: 33
i: [1,42,0,3,1,0,0,1] irreflexivity_r2_xboole_0
s: 42
i: [1,43,0,3,1,0,0,255] reflexivity_r1_tarski
s: 43
i: [1,46,0,3,1,0,0,255] d8_xboole_0
s: 46
i: [1,39,0,6,2,0,0,255] d8_xboole_0
s: 39
d: [2,52,1,3,1,1,40,39,50]
a: [3,56,1,3,1,52]
s: 56
i: [1,44,0,6,2,0,0,1] antisymmetry_r2_xboole_0
s: 44
d: [2,58,1,3,1,1,40,44,50]
a: [3,62,1,3,1,58]
s: 62
i: [1,45,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 45
i: [1,37,0,8,2,0,0,255] d3_tarski
s: 37
d: [2,64,1,8,2,0,40,37,33]
s: 64
i: [1,38,0,8,2,0,0,255] d3_tarski
s: 38
d: [2,68,2,6,2,0,40,38,64]
d: [2,69,2,3,1,0,18,68]
a: [3,74,2,3,1,69]
s: 74
i: [1,36,0,9,3,0,0,255] d3_tarski
s: 36
i: [1,41,0,9,3,0,0,255] d8_xboole_0
s: 41
d: [2,81,3,6,2,1,40,41,74]
d: [2,82,3,3,1,2,21,81,62]
f: 82,69,58,32
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_0__t6_xboole_0
% SZS output start Proof for xboole_0__t6_xboole_0
1. ! [X0] : ! [X1] : ~(! [X2] : ~(~r2_hidden(X2,X0) & r2_hidden(X2,X1)) & r2_xboole_0(X0,X1)) [input t6_xboole_0]
2. ~! [X0] : ! [X1] : ~(! [X2] : ~(~r2_hidden(X2,X0) & r2_hidden(X2,X1)) & r2_xboole_0(X0,X1)) [negated conjecture 1]
4. ! [X0,X1] : (r2_xboole_0(X0,X1) => ~r2_xboole_0(X1,X0)) [input antisymmetry_r2_xboole_0]
5. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
6. ! [X0] : ! [X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [input d8_xboole_0]
10. ~! [X0,X1] : ~(! [X2] : ~(~r2_hidden(X2,X0) & r2_hidden(X2,X1)) & r2_xboole_0(X0,X1)) [flattening 2]
12. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 5]
13. ! [X0,X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [flattening 6]
16. ? [X0,X1] : (! [X2] : (r2_hidden(X2,X0) | ~r2_hidden(X2,X1)) & r2_xboole_0(X0,X1)) [ennf transformation 10]
18. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 12]
19. ! [X0,X1] : (~r2_xboole_0(X1,X0) | ~r2_xboole_0(X0,X1)) [ennf transformation 4]
21. ? [X0,X1] : (! [X2] : (r2_hidden(X2,X0) | ~r2_hidden(X2,X1)) & r2_xboole_0(X0,X1)) => (! [X2] : (r2_hidden(X2,sK0) | ~r2_hidden(X2,sK1)) & r2_xboole_0(sK0,sK1)) [choice axiom]
22. ! [X2] : (r2_hidden(X2,sK0) | ~r2_hidden(X2,sK1)) & r2_xboole_0(sK0,sK1) [skolemisation 16,21]
26. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 18]
27. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 26]
28. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK3(X0,X1),X1) & r2_hidden(sK3(X0,X1),X0))) [choice axiom]
29. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK3(X0,X1),X1) & r2_hidden(sK3(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 27,28]
30. ! [X0,X1] : ((r2_xboole_0(X0,X1) | (X0 = X1 | ~r1_tarski(X0,X1))) & ((X0 != X1 & r1_tarski(X0,X1)) | ~r2_xboole_0(X0,X1))) [nnf transformation 13]
31. ! [X0,X1] : ((r2_xboole_0(X0,X1) | X0 = X1 | ~r1_tarski(X0,X1)) & ((X0 != X1 & r1_tarski(X0,X1)) | ~r2_xboole_0(X0,X1))) [flattening 30]
32. r2_xboole_0(sK0,sK1) [cnf transformation 22] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. ~r2_hidden(X2,sK1) | r2_hidden(X2,sK0) [cnf transformation 22] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. r2_hidden(sK3(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 29] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
38. ~r2_hidden(sK3(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 29] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
41. ~r1_tarski(X0,X1) | X0 = X1 | r2_xboole_0(X0,X1) [cnf transformation 31] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. ~r2_xboole_0(X1,X0) | ~r2_xboole_0(X0,X1) [cnf transformation 19] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
48. 1 <=> r2_xboole_0(sK0,sK1) [avatar definition]
50. r2_xboole_0(sK0,sK1) <- (1) [avatar component clause 48] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
51. 1 [avatar split clause 32,48]
58. ~r2_xboole_0(sK1,sK0) <- (1) [resolution 44,50] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
60. 3 <=> r2_xboole_0(sK1,sK0) [avatar definition]
62. ~r2_xboole_0(sK1,sK0) <- (~3) [avatar component clause 60] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
63. ~3 | ~1 [avatar split clause 58,48,60]
64. r2_hidden(sK3(sK1,X0),sK0) | r1_tarski(sK1,X0) [resolution 37,33] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
68. r1_tarski(sK1,sK0) | r1_tarski(sK1,sK0) [resolution 38,64] {a:2,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
69. r1_tarski(sK1,sK0) [duplicate literal removal 68] {a:2,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
72. 4 <=> r1_tarski(sK1,sK0) [avatar definition]
74. r1_tarski(sK1,sK0) <- (4) [avatar component clause 72] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
75. 4 [avatar split clause 69,72]
81. sK0 = sK1 | r2_xboole_0(sK1,sK0) <- (4) [resolution 41,74] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
82. sK0 = sK1 <- (~3, 4) [subsumption resolution 81,62] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
84. 5 <=> sK0 = sK1 [avatar definition]
87. 5 | 3 | ~4 [avatar split clause 82,72,60,84]
88. sK0 != sK1 | r2_xboole_0(sK1,sK0) | ~r2_xboole_0(sK0,sK1) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
89. $false [avatar sat refutation 51,63,75,87,88] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:31,thDist:-31}
% SZS output end Proof for xboole_0__t6_xboole_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 14
% Generated clauses: 34
% Active clauses: 15
% Passive clauses: 21
% Final active clauses: 15
% Final passive clauses: 5
% Main loop iterations started: 16
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Equational tautologies: 1
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 13
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 8
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

