i: [1,33,0,3,1,1,0,0] t10_arytm_1
a: [3,44,0,3,1,33]
s: 44
i: [1,31,0,3,1,1,0,0] t10_arytm_1
a: [3,54,0,3,1,31]
s: 54
i: [1,30,0,3,1,1,0,0] t10_arytm_1
a: [3,59,0,3,1,30]
s: 59
i: [1,29,0,3,1,1,0,0] t10_arytm_1
a: [3,64,0,3,1,29]
s: 64
i: [1,32,0,5,1,1,0,0] t10_arytm_1
a: [3,49,0,5,1,32]
s: 49
i: [1,34,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,69,0,3,1,34]
s: 69
i: [1,35,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,74,0,3,1,35]
s: 74
i: [1,38,0,4,1,0,0,1] existence_m1_subset_1
s: 38
i: [1,39,0,11,3,0,0,1] dt_k1_arytm_1
s: 39
i: [1,40,0,12,4,0,0,1] connectedness_r1_arytm_2
s: 40
i: [1,37,0,14,4,0,0,1] t9_arytm_1
s: 37
d: [2,76,1,9,3,1,44,39,49]
d: [2,77,1,6,2,2,21,76,64]
d: [2,78,1,3,1,3,21,77,59]
a: [3,82,1,3,1,78]
s: 82
i: [1,36,0,15,5,0,0,1] t4_arytm_1
s: 36
d: [2,88,1,12,4,1,44,37,49]
d: [2,89,1,9,3,1,19,88]
d: [2,90,1,6,2,2,21,89,64]
d: [2,91,1,3,1,3,21,90,59]
d: [2,98,1,12,4,1,40,36,54]
d: [2,99,1,9,3,2,21,98,44]
d: [2,100,1,6,2,3,21,99,64]
d: [2,101,1,3,1,4,21,100,59]
f: 101,91,29,30,31,32,33
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t10_arytm_1
% SZS output start Proof for arytm_1__t10_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((k1_arytm_1(X1,X0) = k11_arytm_3 & r1_arytm_2(X0,X1)) => X0 = X1))) [input t10_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((k1_arytm_1(X1,X0) = k11_arytm_3 & r1_arytm_2(X0,X1)) => X0 = X1))) [negated conjecture 1]
12. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((r1_arytm_2(X1,X0) & r1_arytm_2(X0,X1)) => X0 = X1))) [input t4_arytm_1]
13. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ~(k11_arytm_3 = k1_arytm_1(X0,X1) & ~r1_arytm_2(X0,X1)))) [input t9_arytm_1]
14. ? [X0] : (? [X1] : ((X0 != X1 & (k1_arytm_1(X1,X0) = k11_arytm_3 & r1_arytm_2(X0,X1))) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
15. ? [X0] : (? [X1] : (X0 != X1 & k1_arytm_1(X1,X0) = k11_arytm_3 & r1_arytm_2(X0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 14]
16. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_arytm_2(X1,X0) | ~r1_arytm_2(X0,X1))) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 12]
17. ! [X0] : (! [X1] : (X0 = X1 | ~r1_arytm_2(X1,X0) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 16]
18. ! [X0] : (! [X1] : ((k11_arytm_3 != k1_arytm_1(X0,X1) | r1_arytm_2(X0,X1)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 13]
19. ! [X0] : (! [X1] : (k11_arytm_3 != k1_arytm_1(X0,X1) | r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 18]
24. ? [X0] : (? [X1] : (X0 != X1 & k1_arytm_1(X1,X0) = k11_arytm_3 & r1_arytm_2(X0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (sK0 != X1 & k11_arytm_3 = k1_arytm_1(X1,sK0) & r1_arytm_2(sK0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
25. ? [X1] : (sK0 != X1 & k11_arytm_3 = k1_arytm_1(X1,sK0) & r1_arytm_2(sK0,X1) & m1_subset_1(X1,k2_arytm_2)) => (sK0 != sK1 & k11_arytm_3 = k1_arytm_1(sK1,sK0) & r1_arytm_2(sK0,sK1) & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
26. (sK0 != sK1 & k11_arytm_3 = k1_arytm_1(sK1,sK0) & r1_arytm_2(sK0,sK1) & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 15,25,24]
29. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
30. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
31. r1_arytm_2(sK0,sK1) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
32. k11_arytm_3 = k1_arytm_1(sK1,sK0) [cnf transformation 26] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
33. sK0 != sK1 [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
36. ~r1_arytm_2(X1,X0) | X0 = X1 | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 17] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. k11_arytm_3 != k1_arytm_1(X0,X1) | r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 19] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. 1 <=> sK0 = sK1 [avatar definition]
44. sK0 != sK1 <- (~1) [avatar component clause 42] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
45. ~1 [avatar split clause 33,42]
47. 2 <=> k11_arytm_3 = k1_arytm_1(sK1,sK0) [avatar definition]
49. k11_arytm_3 = k1_arytm_1(sK1,sK0) <- (2) [avatar component clause 47] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
50. 2 [avatar split clause 32,47]
52. 3 <=> r1_arytm_2(sK0,sK1) [avatar definition]
54. r1_arytm_2(sK0,sK1) <- (3) [avatar component clause 52] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
55. 3 [avatar split clause 31,52]
57. 4 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
59. m1_subset_1(sK1,k2_arytm_2) <- (4) [avatar component clause 57] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
60. 4 [avatar split clause 30,57]
62. 5 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
64. m1_subset_1(sK0,k2_arytm_2) <- (5) [avatar component clause 62] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
65. 5 [avatar split clause 29,62]
88. k11_arytm_3 != k11_arytm_3 | r1_arytm_2(sK1,sK0) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (2) [superposition 37,49] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
89. r1_arytm_2(sK1,sK0) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (2) [trivial inequality removal 88] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
90. r1_arytm_2(sK1,sK0) | ~m1_subset_1(sK1,k2_arytm_2) <- (2, 5) [subsumption resolution 89,64] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
91. r1_arytm_2(sK1,sK0) <- (2, 4, 5) [subsumption resolution 90,59] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
93. 9 <=> r1_arytm_2(sK1,sK0) [avatar definition]
96. 9 | ~2 | ~4 | ~5 [avatar split clause 91,62,57,47,93]
98. sK0 = sK1 | ~r1_arytm_2(sK1,sK0) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (3) [resolution 36,54] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
99. ~r1_arytm_2(sK1,sK0) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (~1, 3) [subsumption resolution 98,44] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
100. ~r1_arytm_2(sK1,sK0) | ~m1_subset_1(sK1,k2_arytm_2) <- (~1, 3, 5) [subsumption resolution 99,64] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
101. ~r1_arytm_2(sK1,sK0) <- (~1, 3, 4, 5) [subsumption resolution 100,59] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
102. ~9 | 1 | ~3 | ~4 | ~5 [avatar split clause 101,62,57,52,42,93]
103. $false [avatar sat refutation 45,50,55,60,65,96,102] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:49,thDist:-49}
% SZS output end Proof for arytm_1__t10_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 37
% Active clauses: 13
% Passive clauses: 19
% Final active clauses: 13
% Final passive clauses: 6
% Main loop iterations started: 13
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 7
% 
% >>> Generating Inferences
% Binary resolution: 6
% Forward superposition: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 7
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

