i: [1,101,0,2,1,1,0,0] t12_pre_topc
a: [3,156,0,2,1,101]
s: 156
i: [1,103,0,4,1,1,0,0] t12_pre_topc
a: [3,146,0,4,1,103]
s: 146
i: [1,102,0,5,1,1,0,0] t12_pre_topc
a: [3,151,0,5,1,102]
s: 151
i: [1,104,0,7,2,1,0,0] t12_pre_topc
s: 104
i: [1,105,0,2,1,0,0,255] fc1_xboole_0
a: [3,161,0,2,1,105]
s: 161
i: [1,139,0,2,1,0,0,2] rc2_xboole_0
a: [3,166,0,2,1,139]
s: 166
i: [1,140,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,171,0,2,1,140]
s: 171
i: [1,141,0,2,1,0,0,2] existence_l1_struct_0
a: [3,176,0,2,1,141]
s: 176
i: [1,142,0,2,1,0,0,2] rc1_xboole_0
a: [3,181,0,2,1,142]
s: 181
i: [1,106,0,3,1,0,0,1] fc1_subset_1
s: 106
i: [1,133,0,3,1,0,0,1] rc2_subset_1
s: 133
i: [1,136,0,3,1,0,0,255] reflexivity_r1_tarski
s: 136
i: [1,113,0,4,2,0,0,1] dt_l1_pre_topc
s: 113
d: [2,183,1,2,1,1,40,113,156]
a: [3,188,1,2,1,183]
s: 188
i: [1,131,0,4,1,0,0,1] existence_m1_subset_1
s: 131
i: [1,135,0,4,1,0,0,255] rc3_subset_1
s: 135
d: [2,195,1,5,1,0,40,131,104]
a: [3,199,1,5,1,195]
s: 199
i: [1,110,0,5,2,0,0,1] rc1_subset_1
s: 110
d: [2,184,1,2,1,1,40,113,171]
a: [3,193,1,2,1,184]
s: 193
i: [1,114,0,5,2,0,0,1] fc3_struct_0
s: 114
i: [1,117,0,5,2,0,0,255] t6_boole
s: 117
i: [1,128,0,5,2,0,0,1] t7_boole
s: 128
i: [1,132,0,5,1,0,0,1] rc2_subset_1
s: 132
i: [1,134,0,5,1,0,0,255] rc3_subset_1
s: 134
i: [1,112,0,6,2,0,0,255] rc4_subset_1
s: 112
i: [1,115,0,6,2,0,0,255] d2_struct_0
s: 115
d: [2,203,1,5,2,1,44,146,115]
d: [2,206,1,3,1,2,21,203,188]
a: [3,210,1,3,1,206]
s: 210
i: [1,120,0,6,2,0,0,1] t1_subset
s: 120
i: [1,137,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 137
i: [1,109,0,7,2,0,0,1] rc1_subset_1
s: 109
i: [1,111,0,7,2,0,0,255] rc4_subset_1
s: 111
i: [1,125,0,7,2,0,0,255] t3_subset
s: 125
i: [1,126,0,7,2,0,0,255] t3_subset
s: 126
d: [2,212,1,4,1,1,40,125,151]
a: [3,221,1,4,1,212]
s: 221
i: [1,127,0,7,3,0,0,1] t8_boole
s: 127
i: [1,116,0,8,2,0,0,1] dt_k1_struct_0
s: 116
d: [2,216,1,4,1,0,40,125,132]
s: 216
i: [1,118,0,8,3,0,0,1] cc1_subset_1
s: 118
d: [2,229,1,5,2,1,40,118,151]
a: [3,238,1,3,1,229]
s: 238
i: [1,121,0,8,3,0,0,1] t2_subset
s: 121
i: [1,123,0,8,2,0,0,255] d3_tarski
s: 123
d: [2,246,1,9,2,1,40,121,151]
d: [2,254,1,5,1,1,21,246,106]
a: [3,258,1,5,1,254]
s: 258
i: [1,124,0,8,2,0,0,255] d3_tarski
s: 124
d: [2,267,2,5,1,1,40,258,137]
a: [3,273,2,5,1,267]
s: 273
i: [1,119,0,9,3,0,0,255] cc4_subset_1
s: 119
i: [1,122,0,9,3,0,0,255] d3_tarski
s: 122
d: [2,286,2,7,2,1,40,122,221]
s: 286
i: [1,130,0,9,3,0,0,1] t5_subset
s: 130
d: [2,289,3,7,2,1,40,286,120]
d: [2,292,3,3,1,1,21,289,104]
s: 292
r: 273
r: 199
r: 104
r: 286
i: [1,129,0,10,3,0,0,1] t4_subset
s: 129
i: [1,107,0,11,4,0,0,255] cc3_subset_1
s: 107
a: [3,241,1,2,1,229]
s: 241
d: [2,329,4,0,0,2,52,292,131,241,121]
e: 329
r: 241
i: [1,108,0,11,4,0,0,255] cc2_subset_1
s: 108
a: [3,242,1,2,1,229]
s: 242
d: [2,331,2,0,0,2,52,210,242,117]
e: 331
d: [2,332,2,0,0,3,52,161,210,242,127]
e: 332
d: [2,333,2,0,0,3,52,161,210,242,127]
e: 333
f: 331,329,212,206,183,101,102,103
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pre_topc__t12_pre_topc
% SZS output start Proof for pre_topc__t12_pre_topc
1. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ~(! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ~r2_hidden(X2,X1)) & k1_struct_0(X0) != X1))) [input t12_pre_topc]
2. ~! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ~(! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ~r2_hidden(X2,X1)) & k1_struct_0(X0) != X1))) [negated conjecture 1]
8. ! [X0] : (l1_struct_0(X0) => k1_struct_0(X0) = k1_xboole_0) [input d2_struct_0]
9. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
13. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
20. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
31. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
32. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
33. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
36. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
39. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 31]
40. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 32]
41. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 9]
42. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 33]
48. ? [X0] : (? [X1] : ((! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,u1_struct_0(X0))) & k1_struct_0(X0) != X1) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0)) [ennf transformation 2]
49. ? [X0] : (? [X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,u1_struct_0(X0))) & k1_struct_0(X0) != X1 & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0)) [flattening 48]
56. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 13]
58. ! [X0] : (k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0)) [ennf transformation 8]
60. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 36]
63. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 39]
64. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 40]
65. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 64]
66. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 41]
75. ? [X0] : (? [X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,u1_struct_0(X0))) & k1_struct_0(X0) != X1 & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0)) => (? [X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,u1_struct_0(sK0))) & k1_struct_0(sK0) != X1 & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0)) [choice axiom]
76. ? [X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,u1_struct_0(sK0))) & k1_struct_0(sK0) != X1 & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => (! [X2] : (~r2_hidden(X2,sK1) | ~m1_subset_1(X2,u1_struct_0(sK0))) & k1_struct_0(sK0) != sK1 & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
77. (! [X2] : (~r2_hidden(X2,sK1) | ~m1_subset_1(X2,u1_struct_0(sK0))) & k1_struct_0(sK0) != sK1 & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) [skolemisation 49,76,75]
82. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 66]
83. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 82]
84. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK4(X0,X1),X0))) [choice axiom]
85. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK4(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 83,84]
86. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 42]
87. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK5(X0),X0)) [choice axiom]
88. ! [X0] : m1_subset_1(sK5(X0),X0) [skolemisation 20,87]
101. l1_pre_topc(sK0) [cnf transformation 77] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
102. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 77] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
103. k1_struct_0(sK0) != sK1 [cnf transformation 77] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
104. ~m1_subset_1(X2,u1_struct_0(sK0)) | ~r2_hidden(X2,sK1) [cnf transformation 77] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
113. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 56] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
115. k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0) [cnf transformation 58] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
117. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 60] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
120. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 63] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
121. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 65] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
122. ~r1_tarski(X0,X1) | ~r2_hidden(X3,X0) | r2_hidden(X3,X1) [cnf transformation 85] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
125. ~m1_subset_1(X0,k1_zfmisc_1(X1)) | r1_tarski(X0,X1) [cnf transformation 86] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
131. m1_subset_1(sK5(X0),X0) [cnf transformation 88] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
144. 1 <=> k1_struct_0(sK0) = sK1 [avatar definition]
146. k1_struct_0(sK0) != sK1 <- (~1) [avatar component clause 144] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
147. ~1 [avatar split clause 103,144]
149. 2 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
151. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (2) [avatar component clause 149] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
152. 2 [avatar split clause 102,149]
154. 3 <=> l1_pre_topc(sK0) [avatar definition]
156. l1_pre_topc(sK0) <- (3) [avatar component clause 154] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
157. 3 [avatar split clause 101,154]
183. l1_struct_0(sK0) <- (3) [resolution 113,156] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
186. 9 <=> l1_struct_0(sK0) [avatar definition]
188. l1_struct_0(sK0) <- (9) [avatar component clause 186] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
189. 9 | ~3 [avatar split clause 183,154,186]
203. k1_xboole_0 != sK1 | ~l1_struct_0(sK0) <- (~1) [superposition 146,115] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
206. k1_xboole_0 != sK1 <- (~1, 9) [subsumption resolution 203,188] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
208. 12 <=> k1_xboole_0 = sK1 [avatar definition]
210. k1_xboole_0 != sK1 <- (~12) [avatar component clause 208] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
211. ~12 | 1 | ~9 [avatar split clause 206,186,144,208]
212. r1_tarski(sK1,u1_struct_0(sK0)) <- (2) [resolution 125,151] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
219. 13 <=> r1_tarski(sK1,u1_struct_0(sK0)) [avatar definition]
221. r1_tarski(sK1,u1_struct_0(sK0)) <- (13) [avatar component clause 219] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
222. 13 | ~2 [avatar split clause 212,149,219]
240. 15 <=> v1_xboole_0(sK1) [avatar definition]
241. ~v1_xboole_0(sK1) <- (~15) [avatar component clause 240] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
242. v1_xboole_0(sK1) <- (15) [avatar component clause 240] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
286. r2_hidden(X2,u1_struct_0(sK0)) | ~r2_hidden(X2,sK1) <- (13) [resolution 122,221] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
289. ~r2_hidden(X1,sK1) | m1_subset_1(X1,u1_struct_0(sK0)) <- (13) [resolution 286,120] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
292. ~r2_hidden(X1,sK1) <- (13) [subsumption resolution 289,104] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
329. $false <- (13, ~15) [unit resulting resolution 292,131,241,121] {a:4,w:0,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
330. ~13 | 15 [avatar contradiction clause 329]
331. $false <- (~12, 15) [unit resulting resolution 210,242,117] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
334. 12 | ~15 [avatar contradiction clause 331]
337. $false [avatar sat refutation 147,152,157,189,211,222,330,334] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:127,thDist:-127}
% SZS output end Proof for pre_topc__t12_pre_topc
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 37
% Introduced skolems: 12
% 
% >>> Saturation
% Initial clauses: 42
% Generated clauses: 162
% Active clauses: 53
% Passive clauses: 101
% Final active clauses: 48
% Final passive clauses: 47
% Main loop iterations started: 54
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Fw subsumption resolutions: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 14
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 74
% Unit resulting resolution: 4
% Forward superposition: 1
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 9
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% SAT solver time: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

