i: [1,34,0,20,7,1,0,0] t28_gate_1
a: [3,54,0,2,1,34]
s: 54
a: [3,49,0,8,1,34]
s: 49
i: [1,35,0,2,1,0,0,255] fc1_xboole_0
a: [3,85,0,2,1,35]
s: 85
i: [1,45,0,2,1,0,0,1] rc2_xboole_0
a: [3,90,0,2,1,45]
s: 90
i: [1,46,0,2,1,0,0,1] rc1_xboole_0
a: [3,95,0,2,1,46]
s: 95
i: [1,36,0,5,2,0,0,255] fc2_gate_1
s: 36
i: [1,37,0,5,2,0,0,255] fc1_gate_1
s: 37
i: [1,39,0,11,2,0,0,255] d22_gate_1
s: 39
d: [2,98,1,4,2,1,44,49,39]
d: [2,99,1,2,1,2,21,98,85]
d: [2,100,1,0,0,3,21,99,54]
e: 100
r: 54
a: [3,53,0,2,1,34]
s: 53
a: [3,58,0,2,1,34]
s: 58
i: [1,40,0,11,2,0,0,255] d22_gate_1
s: 40
d: [2,103,1,4,2,1,44,49,40]
d: [2,104,1,2,1,2,21,103,85]
d: [2,105,1,0,0,3,21,104,58]
e: 105
r: 58
a: [3,57,0,2,1,34]
s: 57
a: [3,62,0,2,1,34]
s: 62
i: [1,41,0,11,2,0,0,255] d22_gate_1
s: 41
d: [2,108,1,4,2,1,44,49,41]
d: [2,109,1,2,1,2,21,108,85]
d: [2,110,1,0,0,3,21,109,62]
e: 110
r: 62
a: [3,61,0,2,1,34]
s: 61
a: [3,66,0,2,1,34]
s: 66
i: [1,42,0,11,2,0,0,255] d22_gate_1
s: 42
d: [2,113,1,4,2,1,44,49,42]
d: [2,114,1,2,1,2,21,113,85]
d: [2,115,1,0,0,3,21,114,66]
e: 115
r: 66
a: [3,65,0,2,1,34]
s: 65
a: [3,70,0,2,1,34]
s: 70
i: [1,43,0,11,2,0,0,255] d22_gate_1
s: 43
d: [2,118,1,4,2,1,44,49,43]
d: [2,119,1,2,1,2,21,118,85]
d: [2,120,1,0,0,3,21,119,70]
e: 120
r: 70
a: [3,69,0,2,1,34]
s: 69
a: [3,74,0,2,1,34]
s: 74
i: [1,44,0,11,2,0,0,255] d22_gate_1
s: 44
d: [2,123,1,4,2,1,44,49,44]
d: [2,124,1,2,1,2,21,123,85]
d: [2,125,1,0,0,3,21,124,74]
e: 125
r: 49
r: 74
a: [3,73,0,2,1,34]
s: 73
a: [3,50,0,8,1,34]
s: 50
i: [1,38,0,22,7,0,0,255] d22_gate_1
s: 38
d: [2,145,1,15,7,1,44,50,38]
d: [2,146,1,13,6,2,21,145,73]
d: [2,147,1,11,5,3,21,146,69]
d: [2,148,1,9,4,4,21,147,65]
d: [2,149,1,7,3,5,21,148,61]
d: [2,150,1,5,2,6,21,149,57]
d: [2,151,1,3,1,7,21,150,53]
a: [3,155,1,3,1,151]
s: 155
d: [2,157,2,0,0,2,52,85,155,37]
e: 157
d: [2,158,2,2,1,1,40,155,37]
d: [2,160,2,0,0,2,21,158,85]
e: 160
i: [1,33,0,10,2,1,0,0] t28_gate_1
i: [1,32,0,10,2,1,0,0] t28_gate_1
i: [1,31,0,10,2,1,0,0] t28_gate_1
i: [1,30,0,10,2,1,0,0] t28_gate_1
i: [1,29,0,10,2,1,0,0] t28_gate_1
i: [1,28,0,10,2,1,0,0] t28_gate_1
f: 160,151,125,119,114,109,104,99,35,28,29,30,31,32,33,34
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t28_gate_1
% SZS output start Proof for gate_1__t28_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (~v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5)) <=> (~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input t28_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (~v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5)) <=> (~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ((~(~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_xboole_0) & ~(k22_gate_1(X0,X1,X2,X3,X4,X5) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input d22_gate_1]
14. ~! [X0,X1,X2,X3,X4,X5] : (~v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5)) <=> (~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 2]
15. ! [X0,X1,X2,X3,X4,X5] : ((~(~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_xboole_0) & ~(k22_gate_1(X0,X1,X2,X3,X4,X5) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 13]
16. ? [X0,X1,X2,X3,X4,X5] : (~v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5)) <~> (~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2,X3,X4,X5] : ((k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_xboole_0 | (~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & (k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 15]
20. ? [X0,X1,X2,X3,X4,X5] : (((v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) | v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5))) & ((~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5)))) [nnf transformation 16]
21. ? [X0,X1,X2,X3,X4,X5] : ((v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5))) & ((~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5)))) [flattening 20]
22. ? [X0,X1,X2,X3,X4,X5] : ((v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5))) & ((~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k22_gate_1(X0,X1,X2,X3,X4,X5)))) => ((v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5))) & ((~v1_xboole_0(sK5) & ~v1_xboole_0(sK4) & ~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)))) [choice axiom]
23. (v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5))) & ((~v1_xboole_0(sK5) & ~v1_xboole_0(sK4) & ~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5))) [skolemisation 21,22]
28. ~v1_xboole_0(sK0) | ~v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) [cnf transformation 23] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~v1_xboole_0(sK1) | ~v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) [cnf transformation 23] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. ~v1_xboole_0(sK2) | ~v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) [cnf transformation 23] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~v1_xboole_0(sK3) | ~v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) [cnf transformation 23] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. ~v1_xboole_0(sK4) | ~v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) [cnf transformation 23] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. ~v1_xboole_0(sK5) | ~v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) [cnf transformation 23] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) [cnf transformation 23] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
37. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_xboole_0 | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_xboole_0 | ~v1_xboole_0(X1) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
41. k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_xboole_0 | ~v1_xboole_0(X2) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_xboole_0 | ~v1_xboole_0(X3) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_xboole_0 | ~v1_xboole_0(X4) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. k22_gate_1(X0,X1,X2,X3,X4,X5) = k1_xboole_0 | ~v1_xboole_0(X5) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
48. 1 <=> v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) [avatar definition]
49. ~v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) <- (~1) [avatar component clause 48] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
50. v1_xboole_0(k22_gate_1(sK0,sK1,sK2,sK3,sK4,sK5)) <- (1) [avatar component clause 48] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
52. 2 <=> v1_xboole_0(sK0) [avatar definition]
53. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 52] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. 3 <=> v1_xboole_0(sK1) [avatar definition]
57. ~v1_xboole_0(sK1) <- (~3) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 4 <=> v1_xboole_0(sK2) [avatar definition]
61. ~v1_xboole_0(sK2) <- (~4) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
64. 5 <=> v1_xboole_0(sK3) [avatar definition]
65. ~v1_xboole_0(sK3) <- (~5) [avatar component clause 64] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. 6 <=> v1_xboole_0(sK4) [avatar definition]
69. ~v1_xboole_0(sK4) <- (~6) [avatar component clause 68] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. 7 <=> v1_xboole_0(sK5) [avatar definition]
73. ~v1_xboole_0(sK5) <- (~7) [avatar component clause 72] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
74. v1_xboole_0(sK5) <- (7) [avatar component clause 72] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. 1 | 2 | 3 | 4 | 5 | 6 | 7 [avatar split clause 34,72,68,64,60,56,52,48]
76. ~1 | ~7 [avatar split clause 33,72,48]
77. ~1 | ~6 [avatar split clause 32,68,48]
78. ~1 | ~5 [avatar split clause 31,64,48]
79. ~1 | ~4 [avatar split clause 30,60,48]
80. ~1 | ~3 [avatar split clause 29,56,48]
81. ~1 | ~2 [avatar split clause 28,52,48]
83. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
85. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 83] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
86. 8 [avatar split clause 35,83]
98. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK0) <- (~1) [superposition 49,39] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
99. ~v1_xboole_0(sK0) <- (~1, 8) [subsumption resolution 98,85] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
102. ~2 | 1 | ~8 [avatar split clause 99,83,48,52]
103. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK1) <- (~1) [superposition 49,40] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. ~v1_xboole_0(sK1) <- (~1, 8) [subsumption resolution 103,85] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
107. ~3 | 1 | ~8 [avatar split clause 104,83,48,56]
108. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK2) <- (~1) [superposition 49,41] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
109. ~v1_xboole_0(sK2) <- (~1, 8) [subsumption resolution 108,85] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
112. ~4 | 1 | ~8 [avatar split clause 109,83,48,60]
113. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK3) <- (~1) [superposition 49,42] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. ~v1_xboole_0(sK3) <- (~1, 8) [subsumption resolution 113,85] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
117. ~5 | 1 | ~8 [avatar split clause 114,83,48,64]
118. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK4) <- (~1) [superposition 49,43] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. ~v1_xboole_0(sK4) <- (~1, 8) [subsumption resolution 118,85] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
122. ~6 | 1 | ~8 [avatar split clause 119,83,48,68]
123. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK5) <- (~1) [superposition 49,44] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. ~v1_xboole_0(sK5) <- (~1, 8) [subsumption resolution 123,85] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
125. $false <- (~1, 7, 8) [subsumption resolution 124,74] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. 1 | ~7 | ~8 [avatar contradiction clause 125]
145. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1) [superposition 50,38] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
146. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~7) [subsumption resolution 145,73] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
147. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~6, ~7) [subsumption resolution 146,69] {a:1,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
148. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~5, ~6, ~7) [subsumption resolution 147,65] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
149. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~4, ~5, ~6, ~7) [subsumption resolution 148,61] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
150. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK0) <- (1, ~3, ~4, ~5, ~6, ~7) [subsumption resolution 149,57] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
151. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (1, ~2, ~3, ~4, ~5, ~6, ~7) [subsumption resolution 150,53] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
153. 11 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
155. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (11) [avatar component clause 153] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
156. 11 | ~1 | 2 | 3 | 4 | 5 | 6 | 7 [avatar split clause 151,72,68,64,60,56,52,48,153]
158. ~v1_xboole_0(k1_xboole_0) <- (11) [resolution 155,37] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
160. $false <- (8, 11) [subsumption resolution 158,85] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
161. ~8 | ~11 [avatar contradiction clause 160]
162. $false [avatar sat refutation 75,76,77,78,79,80,81,86,102,107,112,117,122,126,156,161] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:109,thDist:-109}
% SZS output end Proof for gate_1__t28_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 89
% Active clauses: 27
% Passive clauses: 27
% Final active clauses: 20
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 19
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 12
% Backward superposition: 13
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 19
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 3
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

