i: [1,65631,0,49,16,1,0,0] t4_gate_2
a: [3,65656,0,2,1,65631]
s: 65656
a: [3,65665,0,2,1,65631]
s: 65665
a: [3,65673,0,2,1,65631]
s: 65673
a: [3,65681,0,2,1,65631]
s: 65681
a: [3,65688,0,2,1,65631]
s: 65688
a: [3,65697,0,2,1,65631]
s: 65697
a: [3,65705,0,2,1,65631]
s: 65705
a: [3,65712,0,2,1,65631]
s: 65712
i: [1,95,0,15,2,1,0,0] t4_gate_2
a: [3,65977,0,2,1,95]
s: 65977
i: [1,93,0,8,2,1,0,0] t4_gate_2
a: [3,65986,0,2,1,93]
s: 65986
i: [1,91,0,7,2,1,0,0] t4_gate_2
a: [3,65997,0,2,1,91]
s: 65997
i: [1,73,0,7,2,1,0,0] t4_gate_2
a: [3,66054,0,2,1,73]
s: 66054
i: [1,71,0,8,2,1,0,0] t4_gate_2
a: [3,66065,0,2,1,71]
s: 66065
i: [1,69,0,8,2,1,0,0] t4_gate_2
a: [3,66074,0,2,1,69]
s: 66074
i: [1,67,0,9,2,1,0,0] t4_gate_2
a: [3,66084,0,2,1,67]
s: 66084
i: [1,65,0,8,2,1,0,0] t4_gate_2
a: [3,66094,0,2,1,65]
s: 66094
i: [1,63,0,9,2,1,0,0] t4_gate_2
a: [3,66105,0,2,1,63]
s: 66105
i: [1,61,0,9,2,1,0,0] t4_gate_2
a: [3,66115,0,2,1,61]
s: 66115
i: [1,59,0,10,2,1,0,0] t4_gate_2
a: [3,66124,0,2,1,59]
s: 66124
a: [3,65661,0,4,1,65631]
s: 65661
a: [3,65685,0,4,1,65631]
s: 65685
a: [3,65992,0,5,1,91]
s: 65992
i: [1,89,0,7,2,1,0,0] t4_gate_2
a: [3,66003,0,5,1,89]
s: 66003
a: [3,66051,0,5,1,73]
s: 66051
a: [3,65983,0,6,1,93]
s: 65983
i: [1,87,0,8,2,1,0,0] t4_gate_2
a: [3,66009,0,6,1,87]
s: 66009
i: [1,85,0,8,2,1,0,0] t4_gate_2
a: [3,66014,0,6,1,85]
s: 66014
i: [1,81,0,8,2,1,0,0] t4_gate_2
a: [3,66027,0,6,1,81]
s: 66027
a: [3,66060,0,6,1,71]
s: 66060
a: [3,66071,0,6,1,69]
s: 66071
a: [3,66091,0,6,1,65]
s: 66091
i: [1,83,0,9,2,1,0,0] t4_gate_2
a: [3,66021,0,7,1,83]
s: 66021
i: [1,79,0,9,2,1,0,0] t4_gate_2
a: [3,66033,0,7,1,79]
s: 66033
i: [1,77,0,9,2,1,0,0] t4_gate_2
a: [3,66038,0,7,1,77]
s: 66038
a: [3,66081,0,7,1,67]
s: 66081
a: [3,66100,0,7,1,63]
s: 66100
a: [3,66110,0,7,1,61]
s: 66110
i: [1,75,0,10,2,1,0,0] t4_gate_2
a: [3,66044,0,8,1,75]
s: 66044
a: [3,66121,0,8,1,59]
s: 66121
a: [3,65972,0,13,1,95]
s: 65972
i: [1,65632,0,5,2,0,0,1] fc2_gate_1
s: 65632
i: [1,65633,0,5,2,0,0,1] fc1_gate_1
s: 65633
i: [1,65634,0,5,2,0,0,255] t6_boole
s: 65634
i: [1,65639,0,6,2,0,0,1] fc5_gate_1
s: 65639
i: [1,65640,0,6,2,0,0,1] fc4_gate_1
s: 65640
i: [1,65635,0,7,3,0,0,2] t8_boole
s: 65635
d: [2,66130,1,3,1,1,40,65640,65992]
a: [3,66135,1,3,1,66130]
s: 66135
r: 65992
r: 65997
r: 66135
a: [3,65996,0,2,1,91]
s: 65996
a: [3,66134,1,3,1,66130]
s: 66134
a: [3,65993,0,5,1,91]
s: 65993
i: [1,65636,0,7,1,0,0,1] commutativity_k3_gate_1
s: 65636
d: [2,66142,2,2,1,1,40,66135,65632]
a: [3,66145,2,2,1,66142]
s: 66145
i: [1,65637,0,7,1,0,0,1] commutativity_k4_gate_1
s: 65637
d: [2,66131,1,11,1,1,40,65640,65972]
a: [3,66140,1,11,1,66131]
s: 66140
i: [1,65638,0,7,1,0,0,1] commutativity_k2_gate_1
s: 65638
r: 65993
r: 65983
r: 65972
r: 65986
r: 65996
r: 66134
r: 66145
a: [3,65987,0,2,1,93]
s: 65987
s: 65997
d: [2,66152,0,5,1,1,23,65993,65638]
a: [3,66159,0,5,1,66152]
s: 66159
d: [2,66176,1,2,1,1,40,66159,65640]
a: [3,66180,1,2,1,66176]
s: 66180
d: [2,66153,0,6,1,1,23,65983,65638]
a: [3,66164,0,6,1,66153]
s: 66164
d: [2,66154,0,13,1,1,23,65972,65638]
a: [3,66170,0,13,1,66154]
s: 66170
a: [3,66146,2,2,1,66142]
s: 66146
i: [1,65646,0,7,2,0,0,1] fc18_gate_1
s: 65646
d: [2,66185,1,2,1,1,40,65646,66060]
a: [3,66193,1,2,1,66185]
s: 66193
i: [1,65647,0,7,2,0,0,1] fc19_gate_1
s: 65647
i: [1,65648,0,7,2,0,0,1] fc20_gate_1
s: 65648
d: [2,66231,3,0,0,2,52,66146,66110,65648]
e: 66231
d: [2,66235,1,2,1,1,40,65648,66110]
d: [2,66239,1,0,0,2,21,66235,66146]
e: 66239
r: 66110
r: 66115
a: [3,66114,0,2,1,61]
s: 66114
a: [3,66111,0,7,1,61]
s: 66111
d: [2,66208,1,2,1,1,40,65647,66060]
a: [3,66221,1,2,1,66208]
s: 66221
i: [1,65641,0,8,3,0,0,1] fc3_gate_1
s: 65641
i: [1,65642,0,8,3,0,0,1] fc9_gate_1
s: 65642
s: 66135
i: [1,65643,0,8,3,0,0,1] fc7_gate_1
s: 65643
d: [2,66186,1,3,1,1,40,65646,66038]
a: [3,66198,1,3,1,66186]
s: 66198
d: [2,66260,2,0,0,2,52,65977,66198,65632]
e: 66260
d: [2,66261,2,2,1,1,40,66198,65632]
d: [2,66263,2,0,0,2,21,66261,65977]
e: 66263
r: 65656
r: 65712
r: 66038
r: 66044
r: 66198
a: [3,65657,0,2,1,65631]
s: 65657
a: [3,65713,0,2,1,65631]
s: 65713
d: [2,66187,1,3,1,1,40,65646,66110]
a: [3,66203,1,3,1,66187]
s: 66203
d: [2,66265,2,0,0,2,52,66193,66203,65632]
e: 66265
d: [2,66266,2,2,1,1,40,66203,65632]
d: [2,66268,2,0,0,2,21,66266,66193]
e: 66268
r: 66060
r: 66065
r: 66193
a: [3,66064,0,2,1,71]
s: 66064
a: [3,66192,1,2,1,66185]
s: 66192
a: [3,66061,0,6,1,71]
s: 66061
a: [3,66039,0,7,1,77]
s: 66039
d: [2,66207,1,3,1,1,40,65647,66014]
a: [3,66216,1,3,1,66207]
s: 66216
d: [2,66271,2,0,0,2,52,65987,66216,65632]
e: 66271
d: [2,66272,2,2,1,1,40,66216,65632]
d: [2,66274,2,0,0,2,21,66272,65987]
e: 66274
r: 65987
r: 66164
s: 65986
d: [2,66210,1,3,1,1,40,65647,66110]
a: [3,66227,1,3,1,66210]
s: 66227
d: [2,66279,2,0,0,2,52,66221,66227,65632]
e: 66279
d: [2,66280,2,2,1,1,40,66227,65632]
d: [2,66282,2,0,0,2,21,66280,66221]
e: 66282
r: 66100
r: 66105
r: 66221
a: [3,66104,0,2,1,63]
s: 66104
a: [3,66165,0,6,1,66153]
s: 66165
r: 65997
r: 66159
r: 66180
r: 65688
r: 65986
r: 66014
r: 66165
r: 66216
a: [3,65689,0,2,1,65631]
s: 65689
s: 65987
s: 65997
a: [3,65708,0,5,1,65631]
s: 65708
s: 66159
s: 66180
a: [3,66015,0,6,1,85]
s: 66015
a: [3,66220,1,2,1,66208]
s: 66220
s: 66164
a: [3,66101,0,7,1,63]
s: 66101
a: [3,66215,1,3,1,66207]
s: 66215
a: [3,66045,0,8,1,75]
s: 66045
d: [2,66236,1,3,1,1,40,65648,66044]
a: [3,66245,1,3,1,66236]
s: 66245
d: [2,66305,2,0,0,2,52,65997,66245,65632]
e: 66305
d: [2,66306,2,2,1,1,40,66245,65632]
d: [2,66308,2,0,0,2,21,66306,65997]
e: 66308
r: 66245
i: [1,65644,0,8,3,0,0,1] fc6_gate_1
s: 65644
i: [1,65645,0,8,3,0,0,1] fc8_gate_1
s: 65645
d: [2,66312,1,5,2,1,40,65644,65708]
d: [2,66324,1,3,1,2,21,66312,66054]
a: [3,66328,1,3,1,66324]
s: 66328
d: [2,66334,2,0,0,2,52,66180,66328,65632]
e: 66334
d: [2,66335,2,2,1,1,40,66328,65632]
d: [2,66337,2,0,0,2,21,66335,66180]
e: 66337
r: 66051
r: 66054
r: 66328
a: [3,66055,0,2,1,73]
s: 66055
a: [3,66050,0,5,1,73]
s: 66050
d: [2,66340,3,0,0,2,52,66146,66050,65648]
e: 66340
d: [2,66341,2,0,0,2,52,66220,66050,65647]
e: 66341
d: [2,66342,2,0,0,2,52,66192,66050,65646]
e: 66342
d: [2,66343,1,2,1,1,40,66050,65648]
d: [2,66349,1,0,0,2,21,66343,66146]
e: 66349
d: [2,66344,1,2,1,1,40,66050,65647]
d: [2,66351,1,0,0,2,21,66344,66220]
e: 66351
d: [2,66345,1,2,1,1,40,66050,65646]
d: [2,66353,1,0,0,2,21,66345,66192]
e: 66353
r: 65708
r: 66050
r: 66055
s: 66054
a: [3,65676,0,4,1,65631]
s: 65676
d: [2,66355,1,0,0,2,52,66094,65676,65640]
e: 66355
d: [2,66356,1,2,1,1,40,65676,65640]
d: [2,66358,1,0,0,2,21,66356,66094]
e: 66358
r: 65676
a: [3,65652,0,4,1,65631]
s: 65652
d: [2,66360,1,0,0,2,52,66124,65652,65640]
e: 66360
d: [2,66361,1,2,1,1,40,65652,65640]
d: [2,66363,1,0,0,2,21,66361,66124]
e: 66363
r: 65652
a: [3,65700,0,4,1,65631]
s: 65700
d: [2,66365,1,0,0,2,52,66064,65700,65640]
e: 66365
d: [2,66366,1,2,1,1,40,65700,65640]
d: [2,66368,1,0,0,2,21,66366,66064]
e: 66368
r: 65700
a: [3,65692,0,4,1,65631]
s: 65692
d: [2,66370,1,0,0,2,52,66074,65692,65640]
e: 66370
d: [2,66371,1,2,1,1,40,65692,65640]
d: [2,66373,1,0,0,2,21,66371,66074]
e: 66373
r: 65661
r: 65692
a: [3,65660,0,4,1,65631]
s: 65660
d: [2,66375,1,0,0,2,52,66114,65660,65640]
e: 66375
d: [2,66376,1,2,1,1,40,65660,65640]
d: [2,66378,1,0,0,2,21,66376,66114]
e: 66378
r: 65660
r: 65697
r: 66009
a: [3,65696,0,2,1,65631]
s: 65696
a: [3,65693,0,4,1,65631]
s: 65693
a: [3,65709,0,5,1,65631]
s: 65709
a: [3,66327,1,3,1,66324]
s: 66327
s: 66051
a: [3,66008,0,6,1,87]
s: 66008
r: 65696
r: 66008
r: 66121
r: 66124
s: 65697
a: [3,66125,0,2,1,59]
s: 66125
s: 65652
s: 66009
a: [3,66120,0,8,1,59]
s: 66120
d: [2,66311,2,11,2,1,40,65644,66140]
a: [3,66318,2,5,1,66311]
s: 66318
d: [2,66392,3,0,0,2,52,66192,66318,65640]
e: 66392
d: [2,66393,3,2,1,1,40,66318,65640]
d: [2,66395,3,0,0,2,21,66393,66192]
e: 66395
r: 66192
r: 66203
r: 66318
i: [1,65649,0,11,4,0,0,1] fc17_gate_1
s: 65649
d: [2,66399,1,0,0,4,52,65977,65987,65997,66003,65649]
e: 66399
d: [2,66414,1,6,3,1,40,65649,66003]
d: [2,66424,1,4,2,2,21,66414,65997]
d: [2,66425,1,2,1,3,21,66424,65987]
d: [2,66426,1,0,0,4,21,66425,65977]
e: 66426
r: 65705
r: 66003
a: [3,65704,0,2,1,65631]
s: 65704
a: [3,66002,0,5,1,89]
s: 66002
s: 66192
a: [3,66197,1,3,1,66186]
s: 66197
s: 66203
a: [3,66244,1,3,1,66236]
s: 66244
d: [2,66150,1,6,2,0,44,65639,65636]
s: 66150
d: [2,66151,1,6,2,0,44,65639,65636]
s: 66151
a: [3,66317,2,5,1,66311]
s: 66317
d: [2,66155,1,6,2,0,44,65640,65638]
s: 66155
d: [2,66285,1,6,2,1,40,66165,65641]
d: [2,66286,1,4,1,2,21,66285,66180]
a: [3,66289,1,4,1,66286]
s: 66289
d: [2,66453,3,0,0,3,52,66146,66220,66289,65645]
e: 66453
d: [2,66454,2,4,2,1,40,66289,65645]
d: [2,66457,2,2,1,2,21,66454,66220]
d: [2,66458,2,0,0,3,21,66457,66146]
e: 66458
r: 66220
r: 66227
r: 65652
r: 66120
r: 66125
r: 66203
s: 66124
a: [3,65701,0,4,1,65631]
s: 65701
s: 66221
s: 66121
a: [3,66202,1,3,1,66187]
s: 66202
d: [2,66471,2,0,0,2,52,66192,66202,65633]
e: 66471
d: [2,66472,2,2,1,1,40,66202,65633]
d: [2,66474,2,0,0,2,21,66472,66192]
e: 66474
r: 66192
r: 66061
r: 66064
r: 66317
r: 65701
s: 66065
a: [3,65668,0,4,1,65631]
s: 65668
d: [2,66485,1,0,0,2,52,66104,65668,65640]
e: 66485
d: [2,66487,1,2,1,1,40,65668,65640]
d: [2,66489,1,0,0,2,21,66487,66104]
e: 66489
r: 65668
r: 65673
r: 66027
a: [3,65672,0,2,1,65631]
s: 65672
a: [3,65669,0,4,1,65631]
s: 65669
s: 66193
s: 65700
a: [3,66026,0,6,1,81]
s: 66026
d: [2,66494,2,0,0,2,52,66197,66026,65646]
e: 66494
d: [2,66497,1,3,1,1,40,66026,65646]
d: [2,66499,1,0,0,2,21,66497,66197]
e: 66499
r: 65672
r: 65681
r: 66021
r: 66026
s: 65673
a: [3,65680,0,2,1,65631]
s: 65680
a: [3,66226,1,3,1,66210]
s: 66226
a: [3,65677,0,4,1,65631]
s: 65677
s: 66027
s: 66318
s: 66060
a: [3,66322,2,6,1,66311]
s: 66322
d: [2,66509,3,0,0,2,52,66146,66322,65640]
e: 66509
d: [2,66511,3,2,1,1,40,66322,65640]
d: [2,66513,3,0,0,2,21,66511,66146]
e: 66513
r: 66322
a: [3,66020,0,7,1,83]
s: 66020
d: [2,66515,2,0,0,2,52,66244,66020,65648]
e: 66515
d: [2,66516,2,0,0,2,52,66215,66020,65647]
e: 66516
d: [2,66517,1,3,1,1,40,66020,65648]
d: [2,66522,1,0,0,2,21,66517,66244]
e: 66522
d: [2,66518,1,3,1,1,40,66020,65647]
d: [2,66524,1,0,0,2,21,66518,66215]
e: 66524
r: 65680
r: 65685
r: 66020
s: 65681
a: [3,65684,0,4,1,65631]
s: 65684
d: [2,66526,1,0,0,2,52,66084,65684,65640]
e: 66526
d: [2,66528,1,2,1,1,40,65684,65640]
d: [2,66530,1,0,0,2,21,66528,66084]
e: 66530
r: 66081
r: 66084
a: [3,66085,0,2,1,67]
s: 66085
s: 66021
a: [3,66080,0,7,1,67]
s: 66080
d: [2,66534,2,0,0,2,52,66226,66080,65647]
e: 66534
d: [2,66536,1,3,1,1,40,66080,65647]
d: [2,66539,1,0,0,2,21,66536,66226]
e: 66539
r: 65665
r: 65684
r: 66033
r: 66080
r: 66085
a: [3,65664,0,2,1,65631]
s: 65664
s: 66084
s: 65661
s: 65685
a: [3,66032,0,7,1,79]
s: 66032
d: [2,66543,2,0,0,2,52,66244,66032,65648]
e: 66543
d: [2,66544,2,0,0,2,52,66197,66032,65646]
e: 66544
d: [2,66545,1,3,1,1,40,66032,65648]
d: [2,66550,1,0,0,2,21,66545,66244]
e: 66550
d: [2,66547,1,3,1,1,40,66032,65646]
d: [2,66552,1,0,0,2,21,66547,66197]
e: 66552
r: 65664
r: 65704
r: 65977
r: 66002
r: 66032
a: [3,65973,0,13,1,95]
d: [2,66555,0,13,1,1,22,65973,65638]
d: [2,66556,0,0,0,2,21,66555,66170]
e: 66556
r: 65673
r: 66027
r: 66170
r: 66197
s: 65665
s: 65672
s: 65705
a: [3,65976,0,2,1,95]
s: 65976
s: 66003
s: 66026
s: 66198
s: 66033
s: 66081
a: [3,66169,0,13,1,66154]
s: 66169
d: [2,66567,2,0,0,3,52,66180,66140,66169,65641]
e: 66567
d: [2,66568,1,13,2,1,40,66169,65641]
d: [2,66570,1,2,1,2,21,66568,66140]
d: [2,66571,1,0,0,3,21,66570,66180]
e: 66571
r: 66140
a: [3,66139,1,11,1,66131]
s: 66139
d: [2,66575,3,0,0,2,52,66318,66139,66150]
e: 66575
d: [2,66576,2,5,1,1,40,66139,66150]
d: [2,66579,2,0,0,2,21,66576,66318]
e: 66579
r: 65700
r: 66060
r: 66065
r: 66101
r: 66104
r: 66193
r: 66202
r: 66318
r: 65669
r: 65677
r: 66091
r: 66094
s: 66064
a: [3,66095,0,2,1,65]
s: 66095
s: 66105
s: 65668
s: 65676
s: 66192
s: 66061
s: 66203
a: [3,66090,0,6,1,65]
s: 66090
d: [2,66591,3,0,0,2,52,66146,66090,65648]
e: 66591
d: [2,66592,1,2,1,1,40,66090,65648]
d: [2,66596,1,0,0,2,21,66592,66146]
e: 66596
r: 65676
r: 65693
r: 66071
r: 66074
r: 66090
r: 66095
a: [3,66075,0,2,1,69]
s: 66075
s: 66094
s: 65692
s: 66317
a: [3,66070,0,6,1,69]
s: 66070
d: [2,66601,3,0,0,2,52,66146,66070,65648]
e: 66601
d: [2,66602,2,0,0,2,52,66226,66070,65647]
e: 66602
d: [2,66603,2,0,0,2,52,66192,66070,65646]
e: 66603
d: [2,66604,1,2,1,1,40,66070,65648]
d: [2,66610,1,0,0,2,21,66604,66146]
e: 66610
d: [2,66605,1,3,1,1,40,66070,65647]
d: [2,66612,1,0,0,2,21,66605,66226]
e: 66612
d: [2,66606,1,2,1,1,40,66070,65646]
d: [2,66614,1,0,0,2,21,66606,66192]
e: 66614
r: 65692
r: 65705
r: 66003
r: 66070
r: 66075
s: 65704
s: 66074
s: 65701
s: 66002
d: [2,66617,1,0,0,2,52,65976,66002,65646]
e: 66617
d: [2,66620,1,2,1,1,40,66002,65646]
d: [2,66622,1,0,0,2,21,66620,65976]
e: 66622
r: 65672
r: 65704
r: 65987
r: 66002
r: 66026
r: 66164
r: 66289
r: 66215
r: 65657
r: 66039
s: 65656
s: 65673
s: 65705
s: 65986
s: 66216
s: 66003
s: 66027
a: [3,66290,1,4,1,66286]
s: 66290
d: [2,66633,3,0,0,3,52,66146,66221,66290,65643]
e: 66633
d: [2,66634,2,4,2,1,40,66290,65643]
d: [2,66636,2,2,1,2,21,66634,66221]
d: [2,66637,2,0,0,3,21,66636,66146]
e: 66637
r: 65668
r: 66105
r: 66221
r: 66121
r: 66124
r: 66226
s: 66104
s: 66125
a: [3,65653,0,4,1,65631]
s: 65653
d: [2,66642,2,0,0,3,52,66125,66180,65653,65641]
e: 66642
d: [2,66643,1,4,2,1,40,65653,65641]
d: [2,66645,1,2,1,2,21,66643,66180]
d: [2,66646,1,0,0,3,21,66645,66125]
e: 66646
r: 65653
r: 65689
r: 66015
s: 65688
s: 66220
s: 66014
d: [2,66648,1,0,0,2,52,65976,66014,65646]
e: 66648
d: [2,66651,1,2,1,1,40,66014,65646]
d: [2,66653,1,0,0,2,21,66651,65976]
e: 66653
r: 65673
r: 65688
r: 66014
r: 66027
s: 65672
s: 65689
s: 65669
s: 66227
s: 66015
s: 66026
d: [2,66658,1,0,0,2,52,65986,66026,65647]
e: 66658
d: [2,66660,1,2,1,1,40,66026,65647]
d: [2,66663,1,0,0,2,21,66660,65986]
e: 66663
r: 65672
r: 65685
r: 66026
r: 66081
r: 66084
s: 65673
s: 66085
s: 65684
s: 66027
s: 66071
s: 66091
s: 66165
a: [3,66321,2,6,1,66311]
s: 66321
s: 66038
s: 66080
d: [2,66675,2,0,0,2,52,66192,66080,65646]
e: 66675
d: [2,66678,1,2,1,1,40,66080,65646]
d: [2,66680,1,0,0,2,21,66678,66192]
e: 66680
r: 65656
r: 65976
r: 66038
r: 66125
r: 66139
r: 66169
r: 66192
r: 66198
r: 66203
r: 66317
r: 65684
r: 65689
r: 66015
s: 65657
s: 65688
s: 65977
s: 66124
s: 66193
s: 65685
d: [2,66688,2,0,0,3,52,66085,66180,65685,65641]
e: 66688
d: [2,66689,1,4,2,1,40,65685,65641]
d: [2,66691,1,2,1,2,21,66689,66180]
d: [2,66692,1,0,0,3,21,66691,66085]
e: 66692
r: 65685
r: 65705
r: 66003
s: 65704
s: 66197
s: 66002
d: [2,66695,1,0,0,2,52,65986,66002,65647]
e: 66695
d: [2,66697,1,2,1,1,40,66002,65647]
d: [2,66700,1,0,0,2,21,66697,65986]
e: 66700
r: 65688
r: 65704
r: 65997
r: 66002
r: 66159
r: 66244
r: 65681
r: 66021
s: 65680
s: 65689
s: 65705
s: 65996
s: 65684
s: 66202
s: 66003
a: [3,66160,0,5,1,66152]
s: 66160
d: [2,66712,2,0,0,3,52,66180,66135,66160,65641]
e: 66712
d: [2,66713,1,5,2,1,40,66160,65641]
d: [2,66715,1,2,1,2,21,66713,66135]
d: [2,66716,1,0,0,3,21,66715,66180]
e: 66716
r: 66135
r: 66146
r: 65684
r: 66071
r: 66074
r: 66080
r: 66085
s: 66075
s: 66084
s: 66145
s: 65677
s: 65685
s: 66134
s: 66015
s: 66245
s: 66070
s: 66020
d: [2,66671,3,6,2,1,40,66321,65641]
a: [3,66724,3,4,1,66671]
s: 66724
d: [2,66740,4,0,0,3,52,66220,66193,66724,65643]
e: 66740
d: [2,66741,4,4,2,1,40,66724,65643]
d: [2,66743,4,2,1,2,21,66741,66193]
d: [2,66744,4,0,0,3,21,66743,66220]
e: 66744
r: 66321
r: 66724
s: 66039
s: 66318
d: [2,66747,3,0,0,2,52,66134,66318,66155]
e: 66747
d: [2,66748,3,3,1,1,40,66318,66155]
d: [2,66751,3,0,0,2,21,66748,66134]
e: 66751
r: 66318
s: 66081
s: 66101
s: 66322
s: 66121
a: [3,66723,3,4,1,66671]
s: 66723
s: 66170
s: 66140
d: [2,66156,1,6,2,0,44,65640,65638]
s: 66156
d: [2,66129,1,7,3,0,44,65634,65634]
s: 66129
d: [2,66257,1,8,3,0,44,65643,65637]
s: 66257
d: [2,66765,3,0,0,3,52,66220,66145,66290,66257]
e: 66765
d: [2,66768,2,4,2,1,40,66257,66290]
d: [2,66774,2,2,1,2,21,66768,66145]
d: [2,66775,2,0,0,3,21,66774,66220]
e: 66775
r: 66070
r: 66075
r: 66220
r: 66227
r: 66193
r: 66202
r: 66091
r: 66094
r: 65677
r: 65697
r: 66009
s: 65696
s: 66074
s: 66095
s: 65676
s: 66192
s: 65693
s: 66008
d: [2,66795,1,0,0,2,52,65986,66008,65647]
e: 66795
d: [2,66797,1,2,1,1,40,66008,65647]
d: [2,66800,1,0,0,2,21,66797,65986]
e: 66800
r: 65680
r: 65986
r: 66020
r: 66165
r: 66216
r: 66290
s: 65681
s: 65987
s: 66221
s: 66071
s: 66090
s: 66203
s: 66164
s: 66215
s: 66021
s: 66226
s: 66289
d: [2,66817,3,0,0,3,52,66145,66221,66289,65642]
e: 66817
d: [2,66819,2,4,2,1,40,66289,65642]
d: [2,66821,2,2,1,2,21,66819,66221]
d: [2,66822,2,0,0,3,21,66821,66145]
e: 66822
r: 65676
r: 66071
r: 66074
r: 66090
r: 66095
r: 66192
r: 66203
r: 66221
r: 66226
r: 65693
r: 65705
r: 66003
s: 65704
s: 66075
s: 66094
s: 65692
s: 66193
s: 66002
d: [2,66830,1,0,0,2,52,65996,66002,65648]
e: 66830
d: [2,66831,1,2,1,1,40,66002,65648]
d: [2,66835,1,0,0,2,21,66831,65996]
e: 66835
r: 65665
r: 65704
r: 65977
r: 66002
r: 66033
r: 66170
r: 66180
r: 66197
r: 65692
r: 65696
r: 66008
r: 66075
r: 66111
r: 66114
r: 66140
r: 66193
r: 66322
r: 66723
r: 65661
r: 65673
r: 66027
s: 65664
s: 65672
s: 65697
s: 65705
s: 65976
s: 66074
s: 66115
s: 65660
s: 66180
s: 66003
s: 66192
s: 66009
s: 66026
d: [2,66849,1,0,0,2,52,65996,66026,65648]
e: 66849
d: [2,66850,1,2,1,1,40,66026,65648]
d: [2,66854,1,0,0,2,21,66850,65996]
e: 66854
r: 65672
r: 66026
r: 66094
s: 65673
s: 66095
s: 66220
s: 65676
s: 66027
s: 66198
s: 66071
s: 66203
s: 66090
d: [2,66862,2,0,0,2,52,66220,66090,65647]
e: 66862
d: [2,66864,1,2,1,1,40,66090,65647]
d: [2,66867,1,0,0,2,21,66864,66220]
e: 66867
r: 65676
r: 66090
r: 66095
r: 66121
r: 66124
s: 66094
s: 66125
s: 65652
s: 66227
s: 66091
s: 66032
s: 66724
s: 66110
s: 66317
s: 66120
d: [2,66882,2,0,0,2,52,66134,66120,65648]
e: 66882
d: [2,66883,1,3,1,1,40,66120,65648]
d: [2,66887,1,0,0,2,21,66883,66134]
e: 66887
r: 65652
r: 65660
r: 65709
r: 65987
r: 66091
r: 66094
r: 66110
r: 66115
r: 66120
r: 66125
r: 66164
r: 66180
r: 66215
r: 66220
r: 66227
r: 66289
r: 66327
r: 65657
r: 65664
r: 65713
r: 66032
r: 66039
r: 66045
r: 66724
s: 65656
s: 65665
s: 65712
s: 65986
s: 66095
s: 66114
s: 66124
s: 65653
a: [3,66179,1,2,1,66176]
s: 66179
s: 65661
s: 65693
s: 66221
s: 65708
s: 66216
s: 66090
s: 66165
s: 66226
s: 66033
s: 66328
s: 66038
d: [2,66911,1,0,0,2,52,65996,66038,65648]
e: 66911
d: [2,66912,1,2,1,1,40,66038,65648]
d: [2,66916,1,0,0,2,21,66912,65996]
e: 66916
r: 65656
r: 65681
r: 66021
r: 66038
s: 65657
s: 65680
s: 65677
s: 66290
s: 66020
d: [2,66921,1,0,0,2,52,65976,66020,65646]
e: 66921
d: [2,66924,1,2,1,1,40,66020,65646]
d: [2,66926,1,0,0,2,21,66924,65976]
e: 66926
r: 65653
r: 65680
r: 66020
r: 66090
r: 66095
r: 66124
r: 66134
r: 66145
r: 66221
r: 66226
s: 65681
s: 66094
s: 66125
s: 66220
s: 65652
d: [2,66932,2,0,0,2,52,66179,65652,66155]
e: 66932
d: [2,66933,2,2,1,1,40,65652,66155]
d: [2,66936,2,0,0,2,21,66933,66179]
e: 66936
r: 65652
r: 65669
r: 66101
r: 66104
r: 66125
r: 66220
r: 66290
s: 66105
s: 66124
s: 65668
d: [2,66938,2,0,0,2,52,66179,65668,66155]
e: 66938
d: [2,66939,2,2,1,1,40,65668,66155]
d: [2,66942,2,0,0,2,21,66939,66179]
e: 66942
r: 65668
r: 65677
r: 66094
r: 66105
s: 66095
s: 66104
s: 65676
d: [2,66944,2,0,0,2,52,66179,65676,66155]
e: 66944
d: [2,66945,2,2,1,1,40,65676,66155]
d: [2,66948,2,0,0,2,21,66945,66179]
e: 66948
r: 65676
r: 65708
r: 66051
r: 66054
r: 66095
r: 66179
r: 66192
r: 66203
r: 66328
s: 66055
s: 66094
s: 66180
s: 65653
s: 65669
s: 66193
s: 65677
s: 66221
s: 65709
d: [2,66959,2,0,0,2,52,66055,65709,66150]
e: 66959
d: [2,66960,2,2,1,1,40,65709,66150]
d: [2,66963,2,0,0,2,21,66960,66055]
e: 66963
r: 65685
r: 65709
r: 66055
r: 66081
r: 66084
r: 66180
r: 66221
r: 66317
s: 66054
s: 66085
s: 65684
r: 65684
r: 65701
r: 66061
r: 66064
r: 66085
s: 66065
s: 66084
s: 65685
s: 66179
s: 65700
d: [2,66969,2,0,0,2,52,66179,65700,66155]
e: 66969
d: [2,66970,2,2,1,1,40,65700,66155]
d: [2,66973,2,0,0,2,21,66970,66179]
e: 66973
i: [1,49248,0,49,16,1,0,0] t4_gate_2
i: [1,12384,0,49,16,1,0,0] t4_gate_2
i: [1,3168,0,49,16,1,0,0] t4_gate_2
i: [1,864,0,49,16,1,0,0] t4_gate_2
i: [1,288,0,49,16,1,0,0] t4_gate_2
i: [1,144,0,49,16,1,0,0] t4_gate_2
i: [1,108,0,49,16,1,0,0] t4_gate_2
i: [1,99,0,49,16,1,0,0] t4_gate_2
i: [1,96,0,49,16,1,0,0] t4_gate_2
i: [1,94,0,15,2,1,0,0] t4_gate_2
i: [1,92,0,8,2,1,0,0] t4_gate_2
i: [1,90,0,7,2,1,0,0] t4_gate_2
i: [1,88,0,7,2,1,0,0] t4_gate_2
i: [1,86,0,8,2,1,0,0] t4_gate_2
i: [1,84,0,8,2,1,0,0] t4_gate_2
i: [1,82,0,9,2,1,0,0] t4_gate_2
i: [1,80,0,8,2,1,0,0] t4_gate_2
i: [1,78,0,9,2,1,0,0] t4_gate_2
i: [1,76,0,9,2,1,0,0] t4_gate_2
i: [1,74,0,10,2,1,0,0] t4_gate_2
i: [1,72,0,7,2,1,0,0] t4_gate_2
i: [1,70,0,8,2,1,0,0] t4_gate_2
i: [1,68,0,8,2,1,0,0] t4_gate_2
i: [1,66,0,9,2,1,0,0] t4_gate_2
i: [1,64,0,8,2,1,0,0] t4_gate_2
i: [1,62,0,9,2,1,0,0] t4_gate_2
i: [1,60,0,9,2,1,0,0] t4_gate_2
i: [1,58,0,10,2,1,0,0] t4_gate_2
d: [2,66148,2,2,1,1,40,66134,65633]
d: [2,66188,1,3,1,1,40,65646,66044]
d: [2,66189,1,3,1,1,40,65646,66100]
d: [2,66209,1,3,1,1,40,65647,66038]
d: [2,66211,1,3,1,1,40,65647,66044]
d: [2,66212,1,2,1,1,40,65647,66100]
d: [2,66237,1,3,1,1,40,65648,66100]
d: [2,66233,1,3,1,1,40,65648,66060]
d: [2,66232,1,2,1,1,40,65648,66014]
d: [2,66183,1,2,1,1,40,66170,65640]
d: [2,66297,0,5,1,1,22,65992,65638]
d: [2,66301,1,2,1,1,40,66159,65640]
d: [2,66381,1,3,1,1,40,65709,65639]
d: [2,66384,1,3,1,1,40,66008,65648]
d: [2,66391,1,3,1,1,40,66120,65646]
d: [2,66415,1,7,3,1,40,65649,66009]
d: [2,66428,1,5,2,2,21,66415,65987]
d: [2,66429,1,3,1,3,21,66428,65977]
d: [2,66404,1,8,3,1,40,65649,66033]
d: [2,66421,1,6,2,2,21,66404,65987]
d: [2,66403,1,7,3,1,40,65649,66027]
d: [2,66419,1,5,2,2,21,66403,65997]
d: [2,66420,1,3,1,3,21,66419,65987]
d: [2,66460,2,2,1,2,21,66454,66146]
d: [2,66390,1,3,1,1,40,66120,65647]
d: [2,66468,1,9,3,1,40,66121,65649]
d: [2,66469,1,6,2,2,21,66468,66135]
d: [2,66467,1,4,2,1,40,65701,65641]
d: [2,66483,1,2,1,2,21,66467,66180]
d: [2,66437,2,2,1,1,40,66197,65633]
d: [2,66562,1,3,1,1,40,66026,65646]
d: [2,66573,1,11,1,2,21,66568,66180]
d: [2,66491,1,4,2,1,40,65669,65641]
d: [2,66581,1,2,1,2,21,66491,66180]
d: [2,66577,2,6,1,1,40,66139,65639]
d: [2,66304,2,2,1,1,40,66215,65633]
d: [2,66408,1,8,3,1,40,65649,66039]
d: [2,66627,1,6,2,2,21,66408,65997]
d: [2,66628,1,3,1,3,21,66627,66198]
d: [2,66501,2,2,1,1,40,66226,65633]
d: [2,66657,1,7,3,1,40,66015,65649]
d: [2,66682,1,5,2,2,21,66657,65997]
d: [2,66683,1,2,1,3,21,66682,66216]
d: [2,66686,0,13,1,1,22,65972,65638]
d: [2,66439,2,2,1,1,40,66244,65633]
d: [2,66703,0,5,1,1,22,65993,65638]
d: [2,66533,1,8,3,1,40,66021,65649]
d: [2,66705,1,5,2,2,21,66533,66216]
d: [2,66706,1,3,1,3,21,66705,65977]
d: [2,66676,1,3,1,1,40,66080,65648]
d: [2,66755,3,4,1,1,40,66322,66155]
d: [2,66411,1,6,3,1,40,65649,66051]
d: [2,66780,1,4,2,2,21,66411,66145]
d: [2,66760,4,4,2,1,40,66723,65642]
d: [2,66759,4,4,2,1,40,66723,65645]
d: [2,66669,1,7,3,1,40,66091,65649]
d: [2,66786,1,5,2,2,21,66669,66145]
d: [2,66407,1,8,3,1,40,65649,66111]
d: [2,66788,1,6,2,2,21,66407,66145]
d: [2,66729,1,4,2,1,40,65677,65641]
d: [2,66790,1,2,1,2,21,66729,66180]
a: [3,65982,0,6,1,93]
d: [2,66803,0,6,1,1,22,65982,65638]
d: [2,66808,1,3,1,1,40,66090,65646]
d: [2,66811,2,4,1,1,40,66164,66155]
d: [2,66794,1,4,2,1,40,65693,65641]
d: [2,66824,1,2,1,2,21,66794,66180]
d: [2,66826,2,2,1,2,21,66819,66145]
d: [2,66798,1,2,1,1,40,66008,65646]
d: [2,66828,2,2,1,1,40,65692,66155]
d: [2,66812,1,2,1,1,40,66164,65640]
d: [2,66383,2,2,1,1,40,66327,65633]
d: [2,66841,0,13,1,1,22,65973,65638]
d: [2,66541,1,4,2,1,40,65661,65641]
d: [2,66844,2,2,1,1,40,65660,66155]
d: [2,66874,1,2,1,1,40,66032,65647]
d: [2,66409,1,9,3,1,40,65649,66045]
d: [2,66890,1,6,2,2,21,66409,66245]
d: [2,66891,1,3,1,3,21,66890,66198]
d: [2,66876,4,4,2,1,40,66724,66257]
d: [2,66893,4,2,1,2,21,66876,66192]
d: [2,66895,0,6,1,1,22,65983,65638]
d: [2,66900,1,5,2,1,40,65708,65644]
d: [2,66901,1,3,1,2,21,66900,66054]
d: [2,66754,1,8,3,1,40,66101,65649]
d: [2,66928,1,5,2,2,21,66754,66203]
d: [2,66920,2,4,2,1,40,66290,65643]
d: [2,66919,2,4,2,1,40,66290,66257]
d: [2,66907,1,6,2,1,40,66165,65641]
d: [2,66753,1,8,3,1,40,66081,65649]
d: [2,66589,1,7,3,1,40,66061,65649]
d: [2,66881,3,5,2,1,40,66317,65641]
d: [2,66860,1,7,3,1,40,66071,65649]
d: [2,66965,2,2,1,1,40,65684,66155]
f: 66973,66965,66963,66860,66881,66589,66753,66907,66713,66948,66942,66936,66919,66920,66928,66926,66916,66901,66895,66893,66891,66874,66887,66867,66854,66844,66541,66841,66383,66812,66828,66798,66835,66826,66824,66811,66808,66803,66797,66790,66788,66786,66759,66760,66780,66755,66751,66744,66671,66676,66706,66703,66439,66700,66692,66686,66678,66683,66663,66653,66646,66501,66628,66304,66622,66614,66612,66610,66596,66577,66576,66581,66573,66562,66437,66552,66550,66539,66528,66524,66522,66513,66489,66483,66472,66469,66390,66460,66420,66421,66429,66426,66391,66393,66384,66381,66378,66373,66368,66363,66358,66353,66351,66349,66337,66311,66308,66301,66297,66183,66232,66280,66272,66266,66263,66233,66236,66237,66239,66212,66211,66210,66209,66208,66207,66189,66188,66187,66186,66185,66148,66142,66131,66130,58,59,60,61,62,63,64,65,66,67,68,69,70,71,72,73,74,75,76,77,78,79,80,81,82,83,84,85,86,87,88,89,90,91,92,93,94,95,96,99,108,144,288,864,3168,12384,49248
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_2__t4_gate_2
% SZS output start Proof for gate_2__t4_gate_2
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ! [X22] : ~(~(~(v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22)))) & ~(v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X6,X22))) & ~(v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X5,X22))) & ~(v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X4,X22))) & ~(v1_xboole_0(k2_gate_1(X4,X22)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X3,X22))) & ~(v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X2,X22))) & ~(v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X1,X22))) & ~(v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) & ~(v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22))) & ~(v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22)) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22))) & ~(v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [input t4_gate_2]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ! [X22] : ~(~(~(v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22)))) & ~(v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X6,X22))) & ~(v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X5,X22))) & ~(v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X4,X22))) & ~(v1_xboole_0(k2_gate_1(X4,X22)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X3,X22))) & ~(v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X2,X22))) & ~(v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X1,X22))) & ~(v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) & ~(v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22))) & ~(v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22)) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22))) & ~(v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [negated conjecture 1]
4. ! [X0,X1] : k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [input commutativity_k2_gate_1]
5. ! [X0,X1] : k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [input commutativity_k3_gate_1]
6. ! [X0,X1] : k4_gate_1(X0,X1) = k4_gate_1(X1,X0) [input commutativity_k4_gate_1]
13. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc17_gate_1]
14. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc18_gate_1]
15. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X0,X2))) [input fc19_gate_1]
16. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
17. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X2,X0))) [input fc20_gate_1]
18. ! [X0] : (~v1_xboole_0(X0) => v1_xboole_0(k1_gate_1(X0))) [input fc2_gate_1]
19. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_gate_1(X0,X1))) [input fc3_gate_1]
20. ! [X0,X1] : (v1_xboole_0(X0) => v1_xboole_0(k2_gate_1(X0,X1))) [input fc4_gate_1]
21. ! [X0,X1] : (~v1_xboole_0(X1) => ~v1_xboole_0(k3_gate_1(X0,X1))) [input fc5_gate_1]
22. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k3_gate_1(X0,X1))) [input fc6_gate_1]
23. ! [X0,X1] : ((~v1_xboole_0(X1) & v1_xboole_0(X0)) => ~v1_xboole_0(k4_gate_1(X0,X1))) [input fc7_gate_1]
24. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k4_gate_1(X0,X1))) [input fc8_gate_1]
25. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => v1_xboole_0(k4_gate_1(X0,X1))) [input fc9_gate_1]
29. ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22] : ~(~(~(v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22)))) & ~(v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X6,X22))) & ~(v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X5,X22))) & ~(v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X4,X22))) & ~(v1_xboole_0(k2_gate_1(X4,X22)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X3,X22))) & ~(v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X2,X22))) & ~(v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X1,X22))) & ~(v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) & ~(v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22))) & ~(v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22)) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22))) & ~(v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [flattening 2]
34. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22] : (((v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22)))) | (v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22))) & ~v1_xboole_0(X8)) | (v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X6,X22))) | (v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X15)) | (v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X5,X22))) | (v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X14)) | (v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X4,X22))) | (v1_xboole_0(k2_gate_1(X4,X22)) & ~v1_xboole_0(X13)) | (v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X3,X22))) | (v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X12)) | (v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X2,X22))) | (v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X11)) | (v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X1,X22))) | (v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X10)) | (v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) | (v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & (~v1_xboole_0(X21) | v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22))) & (~v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22)) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22))) & (~v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22)) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22)) | v1_xboole_0(X19)) & (~v1_xboole_0(X15) | v1_xboole_0(k8_gate_1(X21,X20,X19))) & (~v1_xboole_0(k8_gate_1(X21,X20,X19)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) | v1_xboole_0(X10)) & (~v1_xboole_0(X9) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) | v1_xboole_0(X9)) & (~v1_xboole_0(X8) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X8)) & (~v1_xboole_0(X7) | v1_xboole_0(k8_gate_1(X18,X17,X16))) & (~v1_xboole_0(k8_gate_1(X18,X17,X16)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X0))) [ennf transformation 29]
35. ! [X0] : (v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0)) [ennf transformation 18]
36. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 16]
39. ! [X0,X1] : (~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1)) [ennf transformation 21]
40. ! [X0,X1] : (v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 20]
41. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 19]
42. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 41]
43. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 25]
44. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 43]
45. ! [X0,X1] : (~v1_xboole_0(k4_gate_1(X0,X1)) | (v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 23]
46. ! [X0,X1] : (~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 45]
47. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 22]
48. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 47]
49. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 24]
50. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 49]
51. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0)) [ennf transformation 14]
52. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0)) [ennf transformation 15]
53. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0)) [ennf transformation 17]
54. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 13]
55. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 54]
56. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22] : (((v1_xboole_0(X8) & ~v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22)))) | (v1_xboole_0(k3_gate_1(X7,k1_gate_1(X22))) & ~v1_xboole_0(X8)) | (v1_xboole_0(X15) & ~v1_xboole_0(k2_gate_1(X6,X22))) | (v1_xboole_0(k2_gate_1(X6,X22)) & ~v1_xboole_0(X15)) | (v1_xboole_0(X14) & ~v1_xboole_0(k2_gate_1(X5,X22))) | (v1_xboole_0(k2_gate_1(X5,X22)) & ~v1_xboole_0(X14)) | (v1_xboole_0(X13) & ~v1_xboole_0(k2_gate_1(X4,X22))) | (v1_xboole_0(k2_gate_1(X4,X22)) & ~v1_xboole_0(X13)) | (v1_xboole_0(X12) & ~v1_xboole_0(k2_gate_1(X3,X22))) | (v1_xboole_0(k2_gate_1(X3,X22)) & ~v1_xboole_0(X12)) | (v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X2,X22))) | (v1_xboole_0(k2_gate_1(X2,X22)) & ~v1_xboole_0(X11)) | (v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(X1,X22))) | (v1_xboole_0(k2_gate_1(X1,X22)) & ~v1_xboole_0(X10)) | (v1_xboole_0(X9) & ~v1_xboole_0(k2_gate_1(X0,X22))) | (v1_xboole_0(k2_gate_1(X0,X22)) & ~v1_xboole_0(X9))) & (~v1_xboole_0(X21) | v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22))) & (~v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))),X22)) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22))) & (~v1_xboole_0(k2_gate_1(k4_gate_1(X16,X17),X22)) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X16),X22)) | v1_xboole_0(X19)) & (~v1_xboole_0(X15) | v1_xboole_0(k8_gate_1(X21,X20,X19))) & (~v1_xboole_0(k8_gate_1(X21,X20,X19)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) | v1_xboole_0(X10)) & (~v1_xboole_0(X9) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) | v1_xboole_0(X9)) & (~v1_xboole_0(X8) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X8)) & (~v1_xboole_0(X7) | v1_xboole_0(k8_gate_1(X18,X17,X16))) & (~v1_xboole_0(k8_gate_1(X18,X17,X16)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X0))) => (((v1_xboole_0(sK8) & ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22)))) | (v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) & ~v1_xboole_0(sK8)) | (v1_xboole_0(sK15) & ~v1_xboole_0(k2_gate_1(sK6,sK22))) | (v1_xboole_0(k2_gate_1(sK6,sK22)) & ~v1_xboole_0(sK15)) | (v1_xboole_0(sK14) & ~v1_xboole_0(k2_gate_1(sK5,sK22))) | (v1_xboole_0(k2_gate_1(sK5,sK22)) & ~v1_xboole_0(sK14)) | (v1_xboole_0(sK13) & ~v1_xboole_0(k2_gate_1(sK4,sK22))) | (v1_xboole_0(k2_gate_1(sK4,sK22)) & ~v1_xboole_0(sK13)) | (v1_xboole_0(sK12) & ~v1_xboole_0(k2_gate_1(sK3,sK22))) | (v1_xboole_0(k2_gate_1(sK3,sK22)) & ~v1_xboole_0(sK12)) | (v1_xboole_0(sK11) & ~v1_xboole_0(k2_gate_1(sK2,sK22))) | (v1_xboole_0(k2_gate_1(sK2,sK22)) & ~v1_xboole_0(sK11)) | (v1_xboole_0(sK10) & ~v1_xboole_0(k2_gate_1(sK1,sK22))) | (v1_xboole_0(k2_gate_1(sK1,sK22)) & ~v1_xboole_0(sK10)) | (v1_xboole_0(sK9) & ~v1_xboole_0(k2_gate_1(sK0,sK22))) | (v1_xboole_0(k2_gate_1(sK0,sK22)) & ~v1_xboole_0(sK9))) & (~v1_xboole_0(sK21) | v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))),sK22))) & (~v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))),sK22)) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(k2_gate_1(k4_gate_1(sK16,sK17),sK22))) & (~v1_xboole_0(k2_gate_1(k4_gate_1(sK16,sK17),sK22)) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k2_gate_1(k1_gate_1(sK16),sK22))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK16),sK22)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0))) [choice axiom]
57. ((v1_xboole_0(sK8) & ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22)))) | (v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) & ~v1_xboole_0(sK8)) | (v1_xboole_0(sK15) & ~v1_xboole_0(k2_gate_1(sK6,sK22))) | (v1_xboole_0(k2_gate_1(sK6,sK22)) & ~v1_xboole_0(sK15)) | (v1_xboole_0(sK14) & ~v1_xboole_0(k2_gate_1(sK5,sK22))) | (v1_xboole_0(k2_gate_1(sK5,sK22)) & ~v1_xboole_0(sK14)) | (v1_xboole_0(sK13) & ~v1_xboole_0(k2_gate_1(sK4,sK22))) | (v1_xboole_0(k2_gate_1(sK4,sK22)) & ~v1_xboole_0(sK13)) | (v1_xboole_0(sK12) & ~v1_xboole_0(k2_gate_1(sK3,sK22))) | (v1_xboole_0(k2_gate_1(sK3,sK22)) & ~v1_xboole_0(sK12)) | (v1_xboole_0(sK11) & ~v1_xboole_0(k2_gate_1(sK2,sK22))) | (v1_xboole_0(k2_gate_1(sK2,sK22)) & ~v1_xboole_0(sK11)) | (v1_xboole_0(sK10) & ~v1_xboole_0(k2_gate_1(sK1,sK22))) | (v1_xboole_0(k2_gate_1(sK1,sK22)) & ~v1_xboole_0(sK10)) | (v1_xboole_0(sK9) & ~v1_xboole_0(k2_gate_1(sK0,sK22))) | (v1_xboole_0(k2_gate_1(sK0,sK22)) & ~v1_xboole_0(sK9))) & (~v1_xboole_0(sK21) | v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))),sK22))) & (~v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))),sK22)) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(k2_gate_1(k4_gate_1(sK16,sK17),sK22))) & (~v1_xboole_0(k2_gate_1(k4_gate_1(sK16,sK17),sK22)) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k2_gate_1(k1_gate_1(sK16),sK22))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK16),sK22)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0)) [skolemisation 34,56]
58. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0) [cnf transformation 57] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. ~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) [cnf transformation 57] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
60. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
61. ~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
62. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
63. ~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
64. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. ~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
66. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
67. ~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
68. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. ~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. ~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
71. ~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
72. ~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7) [cnf transformation 57] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
73. ~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) [cnf transformation 57] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
74. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8) [cnf transformation 57] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
75. ~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) [cnf transformation 57] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
76. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
77. ~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
78. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
79. ~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
80. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
81. ~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
82. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
83. ~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
84. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
85. ~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
86. ~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
87. ~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
88. ~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15) [cnf transformation 57] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
89. ~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) [cnf transformation 57] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
90. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK16),sK22)) | v1_xboole_0(sK19) [cnf transformation 57] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
91. ~v1_xboole_0(sK19) | v1_xboole_0(k2_gate_1(k1_gate_1(sK16),sK22)) [cnf transformation 57] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
92. ~v1_xboole_0(k2_gate_1(k4_gate_1(sK16,sK17),sK22)) | v1_xboole_0(sK20) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
93. ~v1_xboole_0(sK20) | v1_xboole_0(k2_gate_1(k4_gate_1(sK16,sK17),sK22)) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
94. ~v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))),sK22)) | v1_xboole_0(sK21) [cnf transformation 57] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
95. ~v1_xboole_0(sK21) | v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))),sK22)) [cnf transformation 57] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
96. ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 57] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
99. ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK10) | v1_xboole_0(sK9) | v1_xboole_0(k2_gate_1(sK0,sK22)) [cnf transformation 57] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
108. ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK11) | v1_xboole_0(sK10) | v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 57] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
144. ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK12) | v1_xboole_0(sK11) | v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 57] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
288. ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(sK13) | v1_xboole_0(sK12) | v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 57] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
864. ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK14) | v1_xboole_0(sK13) | v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 57] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
3168. ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK15) | v1_xboole_0(sK14) | v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 57] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
12384. ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) | ~v1_xboole_0(sK8) | v1_xboole_0(sK15) | v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 57] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49248. v1_xboole_0(sK8) | v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) | ~v1_xboole_0(k2_gate_1(sK6,sK22)) | ~v1_xboole_0(sK15) | ~v1_xboole_0(k2_gate_1(sK5,sK22)) | ~v1_xboole_0(sK14) | ~v1_xboole_0(k2_gate_1(sK4,sK22)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(k2_gate_1(sK3,sK22)) | ~v1_xboole_0(sK12) | ~v1_xboole_0(k2_gate_1(sK2,sK22)) | ~v1_xboole_0(sK11) | ~v1_xboole_0(k2_gate_1(sK1,sK22)) | ~v1_xboole_0(sK10) | ~v1_xboole_0(k2_gate_1(sK0,sK22)) | ~v1_xboole_0(sK9) [cnf transformation 57] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65632. v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0) [cnf transformation 35] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65633. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 36] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65636. k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [cnf transformation 5] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65637. k4_gate_1(X0,X1) = k4_gate_1(X1,X0) [cnf transformation 6] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65638. k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65639. ~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1) [cnf transformation 39] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65640. v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 40] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65641. ~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 42] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65642. v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 44] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65643. ~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 46] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65644. v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 48] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65645. v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 50] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65646. v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0) [cnf transformation 51] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65647. v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0) [cnf transformation 52] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65648. v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0) [cnf transformation 53] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65649. ~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 55] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65651. 1 <=> v1_xboole_0(k2_gate_1(sK0,sK22)) [avatar definition]
65652. ~v1_xboole_0(k2_gate_1(sK0,sK22)) <- (~1) [avatar component clause 65651] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65653. v1_xboole_0(k2_gate_1(sK0,sK22)) <- (1) [avatar component clause 65651] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65655. 2 <=> v1_xboole_0(sK9) [avatar definition]
65659. 3 <=> v1_xboole_0(k2_gate_1(sK1,sK22)) [avatar definition]
65660. ~v1_xboole_0(k2_gate_1(sK1,sK22)) <- (~3) [avatar component clause 65659] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65661. v1_xboole_0(k2_gate_1(sK1,sK22)) <- (3) [avatar component clause 65659] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65663. 4 <=> v1_xboole_0(sK10) [avatar definition]
65667. 5 <=> v1_xboole_0(k2_gate_1(sK2,sK22)) [avatar definition]
65668. ~v1_xboole_0(k2_gate_1(sK2,sK22)) <- (~5) [avatar component clause 65667] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65669. v1_xboole_0(k2_gate_1(sK2,sK22)) <- (5) [avatar component clause 65667] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65671. 6 <=> v1_xboole_0(sK11) [avatar definition]
65675. 7 <=> v1_xboole_0(k2_gate_1(sK3,sK22)) [avatar definition]
65676. ~v1_xboole_0(k2_gate_1(sK3,sK22)) <- (~7) [avatar component clause 65675] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65677. v1_xboole_0(k2_gate_1(sK3,sK22)) <- (7) [avatar component clause 65675] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65679. 8 <=> v1_xboole_0(sK12) [avatar definition]
65683. 9 <=> v1_xboole_0(k2_gate_1(sK4,sK22)) [avatar definition]
65684. ~v1_xboole_0(k2_gate_1(sK4,sK22)) <- (~9) [avatar component clause 65683] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65685. v1_xboole_0(k2_gate_1(sK4,sK22)) <- (9) [avatar component clause 65683] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65687. 10 <=> v1_xboole_0(sK13) [avatar definition]
65691. 11 <=> v1_xboole_0(k2_gate_1(sK5,sK22)) [avatar definition]
65692. ~v1_xboole_0(k2_gate_1(sK5,sK22)) <- (~11) [avatar component clause 65691] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65693. v1_xboole_0(k2_gate_1(sK5,sK22)) <- (11) [avatar component clause 65691] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65695. 12 <=> v1_xboole_0(sK14) [avatar definition]
65699. 13 <=> v1_xboole_0(k2_gate_1(sK6,sK22)) [avatar definition]
65700. ~v1_xboole_0(k2_gate_1(sK6,sK22)) <- (~13) [avatar component clause 65699] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65701. v1_xboole_0(k2_gate_1(sK6,sK22)) <- (13) [avatar component clause 65699] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65703. 14 <=> v1_xboole_0(sK15) [avatar definition]
65707. 15 <=> v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) [avatar definition]
65708. ~v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) <- (~15) [avatar component clause 65707] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65709. v1_xboole_0(k3_gate_1(sK7,k1_gate_1(sK22))) <- (15) [avatar component clause 65707] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65711. 16 <=> v1_xboole_0(sK8) [avatar definition]
65841. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | 15 | 16 [avatar split clause 49248,65711,65707,65699,65703,65691,65695,65683,65687,65675,65679,65667,65671,65659,65663,65651,65655]
65905. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | 13 | 14 | ~16 | ~15 [avatar split clause 12384,65707,65711,65703,65699,65691,65695,65683,65687,65675,65679,65667,65671,65659,65663,65651,65655]
65937. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | 11 | 12 | ~14 | ~13 | ~16 | ~15 [avatar split clause 3168,65707,65711,65699,65703,65695,65691,65683,65687,65675,65679,65667,65671,65659,65663,65651,65655]
65953. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | 9 | 10 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 864,65707,65711,65699,65703,65691,65695,65687,65683,65675,65679,65667,65671,65659,65663,65651,65655]
65961. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | 7 | 8 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 288,65707,65711,65699,65703,65691,65695,65683,65687,65679,65675,65667,65671,65659,65663,65651,65655]
65965. ~2 | ~1 | ~4 | ~3 | 5 | 6 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 144,65707,65711,65699,65703,65691,65695,65683,65687,65675,65679,65671,65667,65659,65663,65651,65655]
65967. ~2 | ~1 | 3 | 4 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 108,65707,65711,65699,65703,65691,65695,65683,65687,65675,65679,65667,65671,65663,65659,65651,65655]
65968. 1 | 2 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 99,65707,65711,65699,65703,65691,65695,65683,65687,65675,65679,65667,65671,65659,65663,65655,65651]
65969. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 96,65707,65711,65699,65703,65691,65695,65683,65687,65675,65679,65667,65671,65659,65663,65651,65655]
65971. 17 <=> v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))),sK22)) [avatar definition]
65972. ~v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))),sK22)) <- (~17) [avatar component clause 65971] {a:0,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65973. v1_xboole_0(k2_gate_1(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))),sK22)) <- (17) [avatar component clause 65971] {a:0,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65975. 18 <=> v1_xboole_0(sK21) [avatar definition]
65976. v1_xboole_0(sK21) <- (18) [avatar component clause 65975] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65977. ~v1_xboole_0(sK21) <- (~18) [avatar component clause 65975] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65978. 17 | ~18 [avatar split clause 95,65975,65971]
65979. 18 | ~17 [avatar split clause 94,65971,65975]
65981. 19 <=> v1_xboole_0(k2_gate_1(k4_gate_1(sK16,sK17),sK22)) [avatar definition]
65982. ~v1_xboole_0(k2_gate_1(k4_gate_1(sK16,sK17),sK22)) <- (~19) [avatar component clause 65981] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65983. v1_xboole_0(k2_gate_1(k4_gate_1(sK16,sK17),sK22)) <- (19) [avatar component clause 65981] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65985. 20 <=> v1_xboole_0(sK20) [avatar definition]
65986. v1_xboole_0(sK20) <- (20) [avatar component clause 65985] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65987. ~v1_xboole_0(sK20) <- (~20) [avatar component clause 65985] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65988. 19 | ~20 [avatar split clause 93,65985,65981]
65989. 20 | ~19 [avatar split clause 92,65981,65985]
65991. 21 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK16),sK22)) [avatar definition]
65992. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK16),sK22)) <- (~21) [avatar component clause 65991] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65993. v1_xboole_0(k2_gate_1(k1_gate_1(sK16),sK22)) <- (21) [avatar component clause 65991] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65995. 22 <=> v1_xboole_0(sK19) [avatar definition]
65996. v1_xboole_0(sK19) <- (22) [avatar component clause 65995] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65997. ~v1_xboole_0(sK19) <- (~22) [avatar component clause 65995] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65998. 21 | ~22 [avatar split clause 91,65995,65991]
65999. 22 | ~21 [avatar split clause 90,65991,65995]
66001. 23 <=> v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) [avatar definition]
66002. ~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) <- (~23) [avatar component clause 66001] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66003. v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) <- (23) [avatar component clause 66001] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66004. 23 | ~14 [avatar split clause 89,65703,66001]
66005. 14 | ~23 [avatar split clause 88,66001,65703]
66007. 24 <=> v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) [avatar definition]
66008. ~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) <- (~24) [avatar component clause 66007] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66009. v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) <- (24) [avatar component clause 66007] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66010. 24 | ~12 [avatar split clause 87,65695,66007]
66011. 12 | ~24 [avatar split clause 86,66007,65695]
66013. 25 <=> v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) [avatar definition]
66014. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) <- (~25) [avatar component clause 66013] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66015. v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) <- (25) [avatar component clause 66013] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66016. 25 | ~10 [avatar split clause 85,65687,66013]
66017. 10 | ~25 [avatar split clause 84,66013,65687]
66019. 26 <=> v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) [avatar definition]
66020. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) <- (~26) [avatar component clause 66019] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66021. v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) <- (26) [avatar component clause 66019] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66022. 26 | ~8 [avatar split clause 83,65679,66019]
66023. 8 | ~26 [avatar split clause 82,66019,65679]
66025. 27 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) [avatar definition]
66026. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) <- (~27) [avatar component clause 66025] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66027. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) <- (27) [avatar component clause 66025] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66028. 27 | ~6 [avatar split clause 81,65671,66025]
66029. 6 | ~27 [avatar split clause 80,66025,65671]
66031. 28 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) [avatar definition]
66032. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) <- (~28) [avatar component clause 66031] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66033. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) <- (28) [avatar component clause 66031] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66034. 28 | ~4 [avatar split clause 79,65663,66031]
66035. 4 | ~28 [avatar split clause 78,66031,65663]
66037. 29 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) [avatar definition]
66038. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) <- (~29) [avatar component clause 66037] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66039. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) <- (29) [avatar component clause 66037] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66040. 29 | ~2 [avatar split clause 77,65655,66037]
66041. 2 | ~29 [avatar split clause 76,66037,65655]
66043. 30 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) [avatar definition]
66044. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) <- (~30) [avatar component clause 66043] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66045. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) <- (30) [avatar component clause 66043] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66046. 30 | ~16 [avatar split clause 75,65711,66043]
66047. 16 | ~30 [avatar split clause 74,66043,65711]
66049. 31 <=> v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) [avatar definition]
66050. ~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) <- (~31) [avatar component clause 66049] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66051. v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) <- (31) [avatar component clause 66049] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66053. 32 <=> v1_xboole_0(sK7) [avatar definition]
66054. v1_xboole_0(sK7) <- (32) [avatar component clause 66053] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66055. ~v1_xboole_0(sK7) <- (~32) [avatar component clause 66053] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66056. 31 | ~32 [avatar split clause 73,66053,66049]
66057. 32 | ~31 [avatar split clause 72,66049,66053]
66059. 33 <=> v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) [avatar definition]
66060. ~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) <- (~33) [avatar component clause 66059] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66061. v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) <- (33) [avatar component clause 66059] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66063. 34 <=> v1_xboole_0(sK6) [avatar definition]
66064. v1_xboole_0(sK6) <- (34) [avatar component clause 66063] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66066. 33 | ~34 [avatar split clause 71,66063,66059]
66067. 34 | ~33 [avatar split clause 70,66059,66063]
66069. 35 <=> v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) [avatar definition]
66070. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) <- (~35) [avatar component clause 66069] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66071. v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) <- (35) [avatar component clause 66069] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66073. 36 <=> v1_xboole_0(sK5) [avatar definition]
66074. v1_xboole_0(sK5) <- (36) [avatar component clause 66073] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66076. 35 | ~36 [avatar split clause 69,66073,66069]
66077. 36 | ~35 [avatar split clause 68,66069,66073]
66079. 37 <=> v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) [avatar definition]
66080. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) <- (~37) [avatar component clause 66079] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66081. v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) <- (37) [avatar component clause 66079] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66083. 38 <=> v1_xboole_0(sK4) [avatar definition]
66085. ~v1_xboole_0(sK4) <- (~38) [avatar component clause 66083] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66086. 37 | ~38 [avatar split clause 67,66083,66079]
66087. 38 | ~37 [avatar split clause 66,66079,66083]
66089. 39 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) [avatar definition]
66090. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) <- (~39) [avatar component clause 66089] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66091. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) <- (39) [avatar component clause 66089] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66093. 40 <=> v1_xboole_0(sK3) [avatar definition]
66094. v1_xboole_0(sK3) <- (40) [avatar component clause 66093] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66096. 39 | ~40 [avatar split clause 65,66093,66089]
66097. 40 | ~39 [avatar split clause 64,66089,66093]
66099. 41 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) [avatar definition]
66100. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) <- (~41) [avatar component clause 66099] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66101. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) <- (41) [avatar component clause 66099] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66103. 42 <=> v1_xboole_0(sK2) [avatar definition]
66104. v1_xboole_0(sK2) <- (42) [avatar component clause 66103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66106. 41 | ~42 [avatar split clause 63,66103,66099]
66107. 42 | ~41 [avatar split clause 62,66099,66103]
66109. 43 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) [avatar definition]
66110. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) <- (~43) [avatar component clause 66109] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66111. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) <- (43) [avatar component clause 66109] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66113. 44 <=> v1_xboole_0(sK1) [avatar definition]
66114. v1_xboole_0(sK1) <- (44) [avatar component clause 66113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66116. 43 | ~44 [avatar split clause 61,66113,66109]
66117. 44 | ~43 [avatar split clause 60,66109,66113]
66119. 45 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) [avatar definition]
66120. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) <- (~45) [avatar component clause 66119] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66121. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) <- (45) [avatar component clause 66119] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66123. 46 <=> v1_xboole_0(sK0) [avatar definition]
66124. v1_xboole_0(sK0) <- (46) [avatar component clause 66123] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66125. ~v1_xboole_0(sK0) <- (~46) [avatar component clause 66123] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66126. 45 | ~46 [avatar split clause 59,66123,66119]
66127. 46 | ~45 [avatar split clause 58,66119,66123]
66130. ~v1_xboole_0(k1_gate_1(sK16)) <- (~21) [resolution 65640,65992] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66131. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) <- (~17) [resolution 65640,65972] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66133. 47 <=> v1_xboole_0(k1_gate_1(sK16)) [avatar definition]
66134. v1_xboole_0(k1_gate_1(sK16)) <- (47) [avatar component clause 66133] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66135. ~v1_xboole_0(k1_gate_1(sK16)) <- (~47) [avatar component clause 66133] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66136. ~47 | 21 [avatar split clause 66130,65991,66133]
66138. 48 <=> v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) [avatar definition]
66139. v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) <- (48) [avatar component clause 66138] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66140. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) <- (~48) [avatar component clause 66138] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66141. ~48 | 17 [avatar split clause 66131,65971,66138]
66142. v1_xboole_0(sK16) <- (~47) [resolution 66135,65632] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66144. 49 <=> v1_xboole_0(sK16) [avatar definition]
66145. ~v1_xboole_0(sK16) <- (~49) [avatar component clause 66144] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66146. v1_xboole_0(sK16) <- (49) [avatar component clause 66144] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66147. 49 | 47 [avatar split clause 66142,66133,66144]
66148. ~v1_xboole_0(sK16) <- (47) [resolution 66134,65633] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66149. ~49 | ~47 [avatar split clause 66148,66133,66144]
66150. ~v1_xboole_0(k3_gate_1(X1,X0)) | v1_xboole_0(X1) [superposition 65639,65636] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
66155. v1_xboole_0(k2_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 65640,65638] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
66158. 50 <=> v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK16))) [avatar definition]
66159. ~v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK16))) <- (~50) [avatar component clause 66158] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66160. v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK16))) <- (50) [avatar component clause 66158] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66163. 51 <=> v1_xboole_0(k2_gate_1(sK22,k4_gate_1(sK16,sK17))) [avatar definition]
66164. ~v1_xboole_0(k2_gate_1(sK22,k4_gate_1(sK16,sK17))) <- (~51) [avatar component clause 66163] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66165. v1_xboole_0(k2_gate_1(sK22,k4_gate_1(sK16,sK17))) <- (51) [avatar component clause 66163] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66168. 52 <=> v1_xboole_0(k2_gate_1(sK22,k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))))) [avatar definition]
66169. v1_xboole_0(k2_gate_1(sK22,k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))))) <- (52) [avatar component clause 66168] {a:0,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66170. ~v1_xboole_0(k2_gate_1(sK22,k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))))) <- (~52) [avatar component clause 66168] {a:0,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66178. 53 <=> v1_xboole_0(sK22) [avatar definition]
66179. v1_xboole_0(sK22) <- (53) [avatar component clause 66178] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66180. ~v1_xboole_0(sK22) <- (~53) [avatar component clause 66178] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66183. ~v1_xboole_0(sK22) <- (~52) [resolution 66170,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66185. ~v1_xboole_0(sK18) <- (~33) [resolution 65646,66060] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66186. ~v1_xboole_0(k1_gate_1(sK21)) <- (~29) [resolution 65646,66038] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66187. ~v1_xboole_0(k1_gate_1(sK18)) <- (~43) [resolution 65646,66110] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66188. ~v1_xboole_0(k1_gate_1(sK21)) <- (~30) [resolution 65646,66044] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66189. ~v1_xboole_0(k1_gate_1(sK18)) <- (~41) [resolution 65646,66100] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66191. 54 <=> v1_xboole_0(sK18) [avatar definition]
66192. v1_xboole_0(sK18) <- (54) [avatar component clause 66191] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66193. ~v1_xboole_0(sK18) <- (~54) [avatar component clause 66191] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66194. ~54 | 33 [avatar split clause 66185,66059,66191]
66196. 55 <=> v1_xboole_0(k1_gate_1(sK21)) [avatar definition]
66197. v1_xboole_0(k1_gate_1(sK21)) <- (55) [avatar component clause 66196] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66198. ~v1_xboole_0(k1_gate_1(sK21)) <- (~55) [avatar component clause 66196] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66199. ~55 | 29 [avatar split clause 66186,66037,66196]
66201. 56 <=> v1_xboole_0(k1_gate_1(sK18)) [avatar definition]
66202. v1_xboole_0(k1_gate_1(sK18)) <- (56) [avatar component clause 66201] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66203. ~v1_xboole_0(k1_gate_1(sK18)) <- (~56) [avatar component clause 66201] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66204. ~56 | 43 [avatar split clause 66187,66109,66201]
66205. ~55 | 30 [avatar split clause 66188,66043,66196]
66206. ~56 | 41 [avatar split clause 66189,66099,66201]
66207. ~v1_xboole_0(k1_gate_1(sK20)) <- (~25) [resolution 65647,66014] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66208. ~v1_xboole_0(sK17) <- (~33) [resolution 65647,66060] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66209. ~v1_xboole_0(k1_gate_1(sK20)) <- (~29) [resolution 65647,66038] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66210. ~v1_xboole_0(k1_gate_1(sK17)) <- (~43) [resolution 65647,66110] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66211. ~v1_xboole_0(k1_gate_1(sK20)) <- (~30) [resolution 65647,66044] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66212. ~v1_xboole_0(sK17) <- (~41) [resolution 65647,66100] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66214. 57 <=> v1_xboole_0(k1_gate_1(sK20)) [avatar definition]
66215. v1_xboole_0(k1_gate_1(sK20)) <- (57) [avatar component clause 66214] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66216. ~v1_xboole_0(k1_gate_1(sK20)) <- (~57) [avatar component clause 66214] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66217. ~57 | 25 [avatar split clause 66207,66013,66214]
66219. 58 <=> v1_xboole_0(sK17) [avatar definition]
66220. v1_xboole_0(sK17) <- (58) [avatar component clause 66219] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66222. ~58 | 33 [avatar split clause 66208,66059,66219]
66223. ~57 | 29 [avatar split clause 66209,66037,66214]
66225. 59 <=> v1_xboole_0(k1_gate_1(sK17)) [avatar definition]
66226. v1_xboole_0(k1_gate_1(sK17)) <- (59) [avatar component clause 66225] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66227. ~v1_xboole_0(k1_gate_1(sK17)) <- (~59) [avatar component clause 66225] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66228. ~59 | 43 [avatar split clause 66210,66109,66225]
66229. ~57 | 30 [avatar split clause 66211,66043,66214]
66230. ~58 | 41 [avatar split clause 66212,66099,66219]
66232. ~v1_xboole_0(sK19) <- (~25) [resolution 65648,66014] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66233. ~v1_xboole_0(k1_gate_1(sK16)) <- (~33) [resolution 65648,66060] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66235. ~v1_xboole_0(sK16) <- (~43) [resolution 65648,66110] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66236. ~v1_xboole_0(k1_gate_1(sK19)) <- (~30) [resolution 65648,66044] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66237. ~v1_xboole_0(k1_gate_1(sK16)) <- (~41) [resolution 65648,66100] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66239. $false <- (~43, 49) [subsumption resolution 66235,66146] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66240. 43 | ~49 [avatar contradiction clause 66239]
66241. ~47 | 41 [avatar split clause 66237,66099,66133]
66243. 60 <=> v1_xboole_0(k1_gate_1(sK19)) [avatar definition]
66244. v1_xboole_0(k1_gate_1(sK19)) <- (60) [avatar component clause 66243] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66245. ~v1_xboole_0(k1_gate_1(sK19)) <- (~60) [avatar component clause 66243] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66246. ~60 | 30 [avatar split clause 66236,66043,66243]
66247. ~47 | 33 [avatar split clause 66233,66059,66133]
66257. ~v1_xboole_0(k4_gate_1(X1,X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [superposition 65643,65637] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
66261. v1_xboole_0(sK21) <- (~55) [resolution 66198,65632] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66263. $false <- (~18, ~55) [subsumption resolution 66261,65977] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66264. 18 | 55 [avatar contradiction clause 66263]
66266. v1_xboole_0(sK18) <- (~56) [resolution 66203,65632] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66270. 54 | 56 [avatar split clause 66266,66201,66191]
66272. v1_xboole_0(sK20) <- (~57) [resolution 66216,65632] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66276. 20 | 57 [avatar split clause 66272,66214,65985]
66280. v1_xboole_0(sK17) <- (~59) [resolution 66227,65632] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66284. 58 | 59 [avatar split clause 66280,66225,66219]
66288. 61 <=> v1_xboole_0(k4_gate_1(sK16,sK17)) [avatar definition]
66289. ~v1_xboole_0(k4_gate_1(sK16,sK17)) <- (~61) [avatar component clause 66288] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66290. v1_xboole_0(k4_gate_1(sK16,sK17)) <- (61) [avatar component clause 66288] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66292. ~22 | 25 [avatar split clause 66232,66013,65995]
66294. ~53 | 52 [avatar split clause 66183,66168,66178]
66297. ~v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK16))) <- (~21) [forward demodulation 65992,65638] {a:0,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
66298. ~50 | 21 [avatar split clause 66297,65991,66158]
66301. ~v1_xboole_0(sK22) <- (~50) [resolution 66159,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66302. ~53 | 50 [avatar split clause 66301,66158,66178]
66304. ~v1_xboole_0(sK20) <- (57) [resolution 66215,65633] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66306. v1_xboole_0(sK19) <- (~60) [resolution 66245,65632] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66308. $false <- (~22, ~60) [subsumption resolution 66306,65997] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66309. 22 | 60 [avatar contradiction clause 66308]
66311. ~v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) | ~v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) <- (~48) [resolution 65644,66140] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66316. 62 <=> v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) [avatar definition]
66317. v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) <- (62) [avatar component clause 66316] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66318. ~v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) <- (~62) [avatar component clause 66316] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66320. 63 <=> v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) [avatar definition]
66321. v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) <- (63) [avatar component clause 66320] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66322. ~v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) <- (~63) [avatar component clause 66320] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66323. ~62 | ~63 | 48 [avatar split clause 66311,66138,66320,66316]
66326. 64 <=> v1_xboole_0(k1_gate_1(sK22)) [avatar definition]
66327. v1_xboole_0(k1_gate_1(sK22)) <- (64) [avatar component clause 66326] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66328. ~v1_xboole_0(k1_gate_1(sK22)) <- (~64) [avatar component clause 66326] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66335. v1_xboole_0(sK22) <- (~64) [resolution 66328,65632] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66337. $false <- (~53, ~64) [subsumption resolution 66335,66180] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66338. 53 | 64 [avatar contradiction clause 66337]
66343. ~v1_xboole_0(sK16) <- (~31) [resolution 66050,65648] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66344. ~v1_xboole_0(sK17) <- (~31) [resolution 66050,65647] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66345. ~v1_xboole_0(sK18) <- (~31) [resolution 66050,65646] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66349. $false <- (~31, 49) [subsumption resolution 66343,66146] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66350. 31 | ~49 [avatar contradiction clause 66349]
66351. $false <- (~31, 58) [subsumption resolution 66344,66220] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66352. 31 | ~58 [avatar contradiction clause 66351]
66353. $false <- (~31, 54) [subsumption resolution 66345,66192] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66354. 31 | ~54 [avatar contradiction clause 66353]
66356. ~v1_xboole_0(sK3) <- (~7) [resolution 65676,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66358. $false <- (~7, 40) [subsumption resolution 66356,66094] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66359. 7 | ~40 [avatar contradiction clause 66358]
66361. ~v1_xboole_0(sK0) <- (~1) [resolution 65652,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66363. $false <- (~1, 46) [subsumption resolution 66361,66124] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66364. 1 | ~46 [avatar contradiction clause 66363]
66366. ~v1_xboole_0(sK6) <- (~13) [resolution 65700,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66368. $false <- (~13, 34) [subsumption resolution 66366,66064] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66369. 13 | ~34 [avatar contradiction clause 66368]
66371. ~v1_xboole_0(sK5) <- (~11) [resolution 65692,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66373. $false <- (~11, 36) [subsumption resolution 66371,66074] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66374. 11 | ~36 [avatar contradiction clause 66373]
66376. ~v1_xboole_0(sK1) <- (~3) [resolution 65660,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66378. $false <- (~3, 44) [subsumption resolution 66376,66114] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66379. 3 | ~44 [avatar contradiction clause 66378]
66381. v1_xboole_0(k1_gate_1(sK22)) <- (15) [resolution 65709,65639] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66382. 64 | ~15 [avatar split clause 66381,65707,66326]
66383. ~v1_xboole_0(sK22) <- (64) [resolution 66327,65633] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66384. ~v1_xboole_0(k1_gate_1(sK19)) <- (~24) [resolution 66008,65648] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66387. ~60 | 24 [avatar split clause 66384,66007,66243]
66390. ~v1_xboole_0(k1_gate_1(sK17)) <- (~45) [resolution 66120,65647] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66391. ~v1_xboole_0(k1_gate_1(sK18)) <- (~45) [resolution 66120,65646] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66393. ~v1_xboole_0(sK18) <- (~62) [resolution 66318,65640] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66397. ~54 | 62 [avatar split clause 66393,66316,66191]
66398. ~56 | 45 [avatar split clause 66391,66119,66201]
66403. v1_xboole_0(sK19) | v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (27) [resolution 65649,66027] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66404. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (28) [resolution 65649,66033] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66407. v1_xboole_0(sK16) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (43) [resolution 65649,66111] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66408. v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (29) [resolution 65649,66039] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66409. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (30) [resolution 65649,66045] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66411. v1_xboole_0(sK16) | v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (31) [resolution 65649,66051] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66414. v1_xboole_0(sK19) | v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (23) [resolution 65649,66003] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66415. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (24) [resolution 65649,66009] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66419. v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (~22, 27) [subsumption resolution 66403,65997] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66420. v1_xboole_0(k1_gate_1(sK21)) <- (~20, ~22, 27) [subsumption resolution 66419,65987] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66421. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK21)) <- (~20, 28) [subsumption resolution 66404,65987] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66424. v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (~22, 23) [subsumption resolution 66414,65997] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66425. v1_xboole_0(sK21) <- (~20, ~22, 23) [subsumption resolution 66424,65987] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66426. $false <- (~18, ~20, ~22, 23) [subsumption resolution 66425,65977] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66427. 18 | 20 | 22 | ~23 [avatar contradiction clause 66426]
66428. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK21) <- (~20, 24) [subsumption resolution 66415,65987] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66429. v1_xboole_0(k1_gate_1(sK19)) <- (~18, ~20, 24) [subsumption resolution 66428,65977] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66430. 60 | 18 | 20 | ~24 [avatar split clause 66429,66007,65985,65975,66243]
66432. 55 | 60 | 20 | ~28 [avatar split clause 66421,66031,65985,66243,66196]
66433. 55 | 20 | 22 | ~27 [avatar split clause 66420,66025,65995,65985,66196]
66437. ~v1_xboole_0(sK21) <- (55) [resolution 66197,65633] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66439. ~v1_xboole_0(sK19) <- (60) [resolution 66244,65633] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66454. ~v1_xboole_0(sK17) | ~v1_xboole_0(sK16) <- (~61) [resolution 66289,65645] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66460. ~v1_xboole_0(sK17) <- (49, ~61) [subsumption resolution 66454,66146] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66461. ~58 | ~49 | 61 [avatar split clause 66460,66288,66144,66219]
66465. ~59 | 45 [avatar split clause 66390,66119,66225]
66467. v1_xboole_0(sK22) | v1_xboole_0(sK6) <- (13) [resolution 65701,65641] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66468. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (45) [resolution 66121,65649] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66469. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (45, ~47) [subsumption resolution 66468,66135] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66470. 56 | 59 | ~45 | 47 [avatar split clause 66469,66133,66119,66225,66201]
66472. ~v1_xboole_0(sK18) <- (56) [resolution 66202,65633] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66476. ~54 | ~56 [avatar split clause 66472,66201,66191]
66483. v1_xboole_0(sK6) <- (13, ~53) [subsumption resolution 66467,66180] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66484. 34 | ~13 | 53 [avatar split clause 66483,66178,65699,66063]
66487. ~v1_xboole_0(sK2) <- (~5) [resolution 65668,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66489. $false <- (~5, 42) [subsumption resolution 66487,66104] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66490. 5 | ~42 [avatar contradiction clause 66489]
66491. v1_xboole_0(sK22) | v1_xboole_0(sK2) <- (5) [resolution 65669,65641] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66501. ~v1_xboole_0(sK17) <- (59) [resolution 66226,65633] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66511. ~v1_xboole_0(sK16) <- (~63) [resolution 66322,65640] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66513. $false <- (49, ~63) [subsumption resolution 66511,66146] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66514. ~49 | 63 [avatar contradiction clause 66513]
66517. ~v1_xboole_0(k1_gate_1(sK19)) <- (~26) [resolution 66020,65648] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66518. ~v1_xboole_0(k1_gate_1(sK20)) <- (~26) [resolution 66020,65647] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66522. $false <- (~26, 60) [subsumption resolution 66517,66244] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66523. 26 | ~60 [avatar contradiction clause 66522]
66524. $false <- (~26, 57) [subsumption resolution 66518,66215] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66525. 26 | ~57 [avatar contradiction clause 66524]
66528. ~v1_xboole_0(sK4) <- (~9) [resolution 65684,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66532. ~38 | 9 [avatar split clause 66528,65683,66083]
66533. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(sK21) <- (26) [resolution 66021,65649] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66536. ~v1_xboole_0(k1_gate_1(sK17)) <- (~37) [resolution 66080,65647] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66539. $false <- (~37, 59) [subsumption resolution 66536,66226] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66540. 37 | ~59 [avatar contradiction clause 66539]
66541. v1_xboole_0(sK22) | v1_xboole_0(sK1) <- (3) [resolution 65661,65641] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66545. ~v1_xboole_0(k1_gate_1(sK19)) <- (~28) [resolution 66032,65648] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66547. ~v1_xboole_0(k1_gate_1(sK21)) <- (~28) [resolution 66032,65646] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66550. $false <- (~28, 60) [subsumption resolution 66545,66244] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66551. 28 | ~60 [avatar contradiction clause 66550]
66552. $false <- (~28, 55) [subsumption resolution 66547,66197] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66553. 28 | ~55 [avatar contradiction clause 66552]
66554. ~18 | ~55 [avatar split clause 66437,66196,65975]
66562. ~v1_xboole_0(k1_gate_1(sK21)) <- (~27) [resolution 66026,65646] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66563. ~55 | 27 [avatar split clause 66562,66025,66196]
66568. v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) | v1_xboole_0(sK22) <- (52) [resolution 66169,65641] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66573. v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) <- (52, ~53) [subsumption resolution 66568,66180] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66574. 48 | ~52 | 53 [avatar split clause 66573,66178,66168,66138]
66576. v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) <- (48) [resolution 66139,66150] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66577. v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) <- (48) [resolution 66139,65639] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66581. v1_xboole_0(sK2) <- (5, ~53) [subsumption resolution 66491,66180] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66582. 42 | ~5 | 53 [avatar split clause 66581,66178,65667,66103]
66583. 62 | ~48 [avatar split clause 66576,66138,66316]
66584. 63 | ~48 [avatar split clause 66577,66138,66320]
66589. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (33) [resolution 66061,65649] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66592. ~v1_xboole_0(sK16) <- (~39) [resolution 66090,65648] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66596. $false <- (~39, 49) [subsumption resolution 66592,66146] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66597. 39 | ~49 [avatar contradiction clause 66596]
66604. ~v1_xboole_0(sK16) <- (~35) [resolution 66070,65648] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66605. ~v1_xboole_0(k1_gate_1(sK17)) <- (~35) [resolution 66070,65647] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66606. ~v1_xboole_0(sK18) <- (~35) [resolution 66070,65646] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66610. $false <- (~35, 49) [subsumption resolution 66604,66146] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66611. 35 | ~49 [avatar contradiction clause 66610]
66612. $false <- (~35, 59) [subsumption resolution 66605,66226] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66613. 35 | ~59 [avatar contradiction clause 66612]
66614. $false <- (~35, 54) [subsumption resolution 66606,66192] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66615. 35 | ~54 [avatar contradiction clause 66614]
66620. ~v1_xboole_0(sK21) <- (~23) [resolution 66002,65646] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66622. $false <- (18, ~23) [subsumption resolution 66620,65976] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66623. ~18 | 23 [avatar contradiction clause 66622]
66624. ~20 | ~57 [avatar split clause 66304,66214,65985]
66627. v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (~22, 29) [subsumption resolution 66408,65997] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66628. v1_xboole_0(k1_gate_1(sK20)) <- (~22, 29, ~55) [subsumption resolution 66627,66198] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66629. 57 | 22 | ~29 | 55 [avatar split clause 66628,66196,66037,65995,66214]
66641. ~58 | ~59 [avatar split clause 66501,66225,66219]
66643. v1_xboole_0(sK22) | v1_xboole_0(sK0) <- (1) [resolution 65653,65641] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66645. v1_xboole_0(sK0) <- (1, ~53) [subsumption resolution 66643,66180] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66646. $false <- (1, ~46, ~53) [subsumption resolution 66645,66125] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66647. ~1 | 46 | 53 [avatar contradiction clause 66646]
66651. ~v1_xboole_0(sK21) <- (~25) [resolution 66014,65646] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66653. $false <- (18, ~25) [subsumption resolution 66651,65976] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66654. ~18 | 25 [avatar contradiction clause 66653]
66657. v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(sK21) <- (25) [resolution 66015,65649] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66660. ~v1_xboole_0(sK20) <- (~27) [resolution 66026,65647] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66663. $false <- (20, ~27) [subsumption resolution 66660,65986] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66664. ~20 | 27 [avatar contradiction clause 66663]
66669. v1_xboole_0(sK16) | v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (39) [resolution 66091,65649] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66671. v1_xboole_0(k4_gate_1(sK17,sK18)) | v1_xboole_0(sK16) <- (63) [resolution 66321,65641] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66676. ~v1_xboole_0(k1_gate_1(sK16)) <- (~37) [resolution 66080,65648] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66678. ~v1_xboole_0(sK18) <- (~37) [resolution 66080,65646] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66682. v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(sK21) <- (~22, 25) [subsumption resolution 66657,65997] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66683. v1_xboole_0(sK21) <- (~22, 25, ~57) [subsumption resolution 66682,66216] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66684. 18 | 22 | ~25 | 57 [avatar split clause 66683,66214,66013,65995,65975]
66685. ~54 | 37 [avatar split clause 66678,66079,66191]
66686. ~v1_xboole_0(k2_gate_1(sK22,k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))))) <- (~17) [forward demodulation 65972,65638] {a:0,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
66687. ~52 | 17 [avatar split clause 66686,65971,66168]
66689. v1_xboole_0(sK22) | v1_xboole_0(sK4) <- (9) [resolution 65685,65641] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66691. v1_xboole_0(sK4) <- (9, ~53) [subsumption resolution 66689,66180] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66692. $false <- (9, ~38, ~53) [subsumption resolution 66691,66085] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66693. ~9 | 38 | 53 [avatar contradiction clause 66692]
66697. ~v1_xboole_0(sK20) <- (~23) [resolution 66002,65647] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66700. $false <- (20, ~23) [subsumption resolution 66697,65986] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66701. ~20 | 23 [avatar contradiction clause 66700]
66702. ~22 | ~60 [avatar split clause 66439,66243,65995]
66703. v1_xboole_0(k2_gate_1(sK22,k1_gate_1(sK16))) <- (21) [forward demodulation 65993,65638] {a:0,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
66704. 50 | ~21 [avatar split clause 66703,65991,66158]
66705. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK21) <- (26, ~57) [subsumption resolution 66533,66216] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66706. v1_xboole_0(k1_gate_1(sK19)) <- (~18, 26, ~57) [subsumption resolution 66705,65977] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66707. 60 | 18 | ~26 | 57 [avatar split clause 66706,66214,66019,65975,66243]
66713. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK22) <- (50) [resolution 66160,65641] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66720. ~47 | 37 [avatar split clause 66676,66079,66133]
66722. 65 <=> v1_xboole_0(k4_gate_1(sK17,sK18)) [avatar definition]
66723. ~v1_xboole_0(k4_gate_1(sK17,sK18)) <- (~65) [avatar component clause 66722] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
66724. v1_xboole_0(k4_gate_1(sK17,sK18)) <- (65) [avatar component clause 66722] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
66725. 49 | 65 | ~63 [avatar split clause 66671,66320,66722,66144]
66729. v1_xboole_0(sK22) | v1_xboole_0(sK3) <- (7) [resolution 65677,65641] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66741. v1_xboole_0(sK18) | ~v1_xboole_0(sK17) <- (65) [resolution 66724,65643] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
66743. ~v1_xboole_0(sK17) <- (~54, 65) [subsumption resolution 66741,66193] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
66744. $false <- (~54, 58, 65) [subsumption resolution 66743,66220] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
66745. 54 | ~58 | ~65 [avatar contradiction clause 66744]
66748. ~v1_xboole_0(k1_gate_1(sK16)) <- (~62) [resolution 66318,66155] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
66751. $false <- (47, ~62) [subsumption resolution 66748,66134] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
66752. ~47 | 62 [avatar contradiction clause 66751]
66753. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (37) [resolution 66081,65649] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66754. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (41) [resolution 66101,65649] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66755. ~v1_xboole_0(k4_gate_1(sK17,sK18)) <- (~63) [resolution 66322,66155] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
66757. ~65 | 63 [avatar split clause 66755,66320,66722]
66759. ~v1_xboole_0(sK18) | ~v1_xboole_0(sK17) <- (~65) [resolution 66723,65645] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
66760. v1_xboole_0(sK18) | v1_xboole_0(sK17) <- (~65) [resolution 66723,65642] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
66780. v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (31, ~49) [subsumption resolution 66411,66145] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66783. 54 | 58 | ~31 | 49 [avatar split clause 66780,66144,66049,66219,66191]
66784. 58 | 54 | 65 [avatar split clause 66760,66722,66191,66219]
66785. ~58 | ~54 | 65 [avatar split clause 66759,66722,66191,66219]
66786. v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (39, ~49) [subsumption resolution 66669,66145] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66787. 56 | 58 | ~39 | 49 [avatar split clause 66786,66144,66089,66219,66201]
66788. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (43, ~49) [subsumption resolution 66407,66145] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66789. 56 | 59 | ~43 | 49 [avatar split clause 66788,66144,66109,66225,66201]
66790. v1_xboole_0(sK3) <- (7, ~53) [subsumption resolution 66729,66180] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66791. 40 | ~7 | 53 [avatar split clause 66790,66178,65675,66093]
66794. v1_xboole_0(sK22) | v1_xboole_0(sK5) <- (11) [resolution 65693,65641] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66797. ~v1_xboole_0(sK20) <- (~24) [resolution 66008,65647] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66798. ~v1_xboole_0(sK21) <- (~24) [resolution 66008,65646] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66802. ~20 | 24 [avatar split clause 66797,66007,65985]
66803. ~v1_xboole_0(k2_gate_1(sK22,k4_gate_1(sK16,sK17))) <- (~19) [forward demodulation 65982,65638] {a:0,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
66804. ~51 | 19 [avatar split clause 66803,65981,66163]
66808. ~v1_xboole_0(k1_gate_1(sK18)) <- (~39) [resolution 66090,65646] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66809. ~56 | 39 [avatar split clause 66808,66089,66201]
66811. ~v1_xboole_0(k4_gate_1(sK16,sK17)) <- (~51) [resolution 66164,66155] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66812. ~v1_xboole_0(sK22) <- (~51) [resolution 66164,65640] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66813. ~61 | 51 [avatar split clause 66811,66163,66288]
66819. v1_xboole_0(sK17) | v1_xboole_0(sK16) <- (~61) [resolution 66289,65642] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66824. v1_xboole_0(sK5) <- (11, ~53) [subsumption resolution 66794,66180] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66825. 36 | ~11 | 53 [avatar split clause 66824,66178,65691,66073]
66826. v1_xboole_0(sK17) <- (~49, ~61) [subsumption resolution 66819,66145] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66827. 58 | 49 | 61 [avatar split clause 66826,66288,66144,66219]
66828. ~v1_xboole_0(sK22) <- (~11) [resolution 65692,66155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66831. ~v1_xboole_0(sK19) <- (~23) [resolution 66002,65648] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66835. $false <- (22, ~23) [subsumption resolution 66831,65996] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66836. ~22 | 23 [avatar contradiction clause 66835]
66837. ~18 | 24 [avatar split clause 66798,66007,65975]
66838. ~53 | 11 [avatar split clause 66828,65691,66178]
66839. ~53 | 51 [avatar split clause 66812,66163,66178]
66840. ~53 | ~64 [avatar split clause 66383,66326,66178]
66841. v1_xboole_0(k2_gate_1(sK22,k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))))) <- (17) [forward demodulation 65973,65638] {a:0,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
66842. 52 | ~17 [avatar split clause 66841,65971,66168]
66843. 44 | 53 | ~3 [avatar split clause 66541,65659,66178,66113]
66844. ~v1_xboole_0(sK22) <- (~3) [resolution 65660,66155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66846. ~53 | 3 [avatar split clause 66844,65659,66178]
66850. ~v1_xboole_0(sK19) <- (~27) [resolution 66026,65648] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66854. $false <- (22, ~27) [subsumption resolution 66850,65996] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66855. ~22 | 27 [avatar contradiction clause 66854]
66860. v1_xboole_0(sK16) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (35) [resolution 66071,65649] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66864. ~v1_xboole_0(sK17) <- (~39) [resolution 66090,65647] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66867. $false <- (~39, 58) [subsumption resolution 66864,66220] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66868. 39 | ~58 [avatar contradiction clause 66867]
66874. ~v1_xboole_0(sK20) <- (~28) [resolution 66032,65647] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66876. v1_xboole_0(sK17) | ~v1_xboole_0(sK18) <- (65) [resolution 66724,66257] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
66881. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK18) <- (62) [resolution 66317,65641] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66883. ~v1_xboole_0(k1_gate_1(sK16)) <- (~45) [resolution 66120,65648] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66887. $false <- (~45, 47) [subsumption resolution 66883,66134] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66888. 45 | ~47 [avatar contradiction clause 66887]
66889. ~20 | 28 [avatar split clause 66874,66031,65985]
66890. v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (30, ~60) [subsumption resolution 66409,66245] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66891. v1_xboole_0(k1_gate_1(sK20)) <- (30, ~55, ~60) [subsumption resolution 66890,66198] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66892. 57 | ~30 | 55 | 60 [avatar split clause 66891,66243,66196,66043,66214]
66893. v1_xboole_0(sK17) <- (54, 65) [subsumption resolution 66876,66192] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
66894. 58 | ~54 | ~65 [avatar split clause 66893,66722,66191,66219]
66895. v1_xboole_0(k2_gate_1(sK22,k4_gate_1(sK16,sK17))) <- (19) [forward demodulation 65983,65638] {a:0,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
66896. 51 | ~19 [avatar split clause 66895,65981,66163]
66900. ~v1_xboole_0(k1_gate_1(sK22)) | ~v1_xboole_0(sK7) <- (~15) [resolution 65708,65644] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66901. ~v1_xboole_0(k1_gate_1(sK22)) <- (~15, 32) [subsumption resolution 66900,66054] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66902. ~64 | 15 | ~32 [avatar split clause 66901,66053,65707,66326]
66907. v1_xboole_0(k4_gate_1(sK16,sK17)) | v1_xboole_0(sK22) <- (51) [resolution 66165,65641] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66912. ~v1_xboole_0(sK19) <- (~29) [resolution 66038,65648] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66916. $false <- (22, ~29) [subsumption resolution 66912,65996] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66917. ~22 | 29 [avatar contradiction clause 66916]
66919. v1_xboole_0(sK16) | ~v1_xboole_0(sK17) <- (61) [resolution 66290,66257] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66920. v1_xboole_0(sK17) | ~v1_xboole_0(sK16) <- (61) [resolution 66290,65643] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66924. ~v1_xboole_0(sK21) <- (~26) [resolution 66020,65646] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66926. $false <- (18, ~26) [subsumption resolution 66924,65976] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66927. ~18 | 26 [avatar contradiction clause 66926]
66928. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK17) <- (41, ~56) [subsumption resolution 66754,66203] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66929. 58 | 47 | ~41 | 56 [avatar split clause 66928,66201,66099,66133,66219]
66930. ~49 | 58 | ~61 [avatar split clause 66920,66288,66219,66144]
66931. ~58 | 49 | ~61 [avatar split clause 66919,66288,66144,66219]
66933. ~v1_xboole_0(sK22) <- (~1) [resolution 65652,66155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66936. $false <- (~1, 53) [subsumption resolution 66933,66179] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66937. 1 | ~53 [avatar contradiction clause 66936]
66939. ~v1_xboole_0(sK22) <- (~5) [resolution 65668,66155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66942. $false <- (~5, 53) [subsumption resolution 66939,66179] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66943. 5 | ~53 [avatar contradiction clause 66942]
66945. ~v1_xboole_0(sK22) <- (~7) [resolution 65676,66155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66948. $false <- (~7, 53) [subsumption resolution 66945,66179] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66949. 7 | ~53 [avatar contradiction clause 66948]
66950. 53 | 47 | ~50 [avatar split clause 66713,66158,66133,66178]
66951. 53 | 61 | ~51 [avatar split clause 66907,66163,66288,66178]
66952. 54 | 59 | 47 | ~37 [avatar split clause 66753,66079,66133,66225,66191]
66953. 54 | 58 | 47 | ~33 [avatar split clause 66589,66059,66133,66219,66191]
66954. 54 | 47 | ~62 [avatar split clause 66881,66316,66133,66191]
66955. 54 | 59 | 49 | ~35 [avatar split clause 66860,66069,66144,66225,66191]
66960. v1_xboole_0(sK7) <- (15) [resolution 65709,66150] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66963. $false <- (15, ~32) [subsumption resolution 66960,66055] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66964. ~15 | 32 [avatar contradiction clause 66963]
66965. ~v1_xboole_0(sK22) <- (~9) [resolution 65684,66155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66967. ~53 | 9 [avatar split clause 66965,65683,66178]
66970. ~v1_xboole_0(sK22) <- (~13) [resolution 65700,66155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66973. $false <- (~13, 53) [subsumption resolution 66970,66179] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66974. 13 | ~53 [avatar contradiction clause 66973]
66975. $false [avatar sat refutation 65841,65905,65937,65953,65961,65965,65967,65968,65969,65978,65979,65988,65989,65998,65999,66004,66005,66010,66011,66016,66017,66022,66023,66028,66029,66034,66035,66040,66041,66046,66047,66056,66057,66066,66067,66076,66077,66086,66087,66096,66097,66106,66107,66116,66117,66126,66127,66136,66141,66147,66149,66194,66199,66204,66205,66206,66217,66222,66223,66228,66229,66230,66240,66241,66246,66247,66264,66270,66276,66284,66292,66294,66298,66302,66309,66323,66338,66350,66352,66354,66359,66364,66369,66374,66379,66382,66387,66397,66398,66427,66430,66432,66433,66461,66465,66470,66476,66484,66490,66514,66523,66525,66532,66540,66551,66553,66554,66563,66574,66582,66583,66584,66597,66611,66613,66615,66623,66624,66629,66641,66647,66654,66664,66684,66685,66687,66693,66701,66702,66704,66707,66720,66725,66745,66752,66757,66783,66784,66785,66787,66789,66791,66802,66804,66809,66813,66825,66827,66836,66837,66838,66839,66840,66842,66843,66846,66855,66868,66888,66889,66892,66894,66896,66902,66917,66927,66929,66930,66931,66937,66943,66949,66950,66951,66952,66953,66954,66955,66964,66967,66974] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:6078,thDist:-6078}
% SZS output end Proof for gate_2__t4_gate_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 27
% >>> Preprocessing
% Introduced skolems: 23
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 65592
% Generated clauses: 66678
% Active clauses: 370
% Passive clauses: 461
% Final active clauses: 60
% Final passive clauses: 24
% Main loop iterations started: 373
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 113
% Fw demodulations: 13
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 65292
% Forward subsumptions: 254
% 
% >>> Generating Inferences
% Binary resolution: 319
% Unit resulting resolution: 59
% Forward superposition: 16
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 317
% Split components: 4221
% Sat splitting refutations: 119
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 535
% SAT solver binary clauses: 207
% 
% Memory used [KB]: 31854
% Time elapsed: 0.116 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.049 s
% preprocessing: 0.010 s ( own 0.009 s ) 
% minimizing solver time: 0.002 s
% sat proof minimization: 0.001 s
% term sharing: 0.002 s
% congruence closure: 0.003 s ( own 0.001 s ) 
% forward subsumption: 0.001 s
% unit resulting resolution: 0.005 s ( own 0.004 s ) 
% superposition: 0.001 s
% other: 0.116 s ( own 0.046 s ) 

