i: [1,45,0,10,2,1,0,0] t59_zfmisc_1
a: [3,69,0,3,1,45]
s: 69
a: [3,66,0,7,1,45]
s: 66
i: [1,46,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,76,0,2,1,46]
s: 76
i: [1,47,0,2,1,0,0,2] fc1_xboole_0
a: [3,81,0,2,1,47]
s: 81
i: [1,61,0,2,1,0,0,2] rc2_xboole_0
a: [3,91,0,2,1,61]
s: 91
i: [1,62,0,2,1,0,0,2] rc1_xboole_0
a: [3,96,0,2,1,62]
s: 96
i: [1,49,0,3,1,0,0,1] fc2_xboole_0
s: 49
i: [1,48,0,3,1,0,0,255] d2_xboole_0
a: [3,86,0,3,1,48]
s: 86
r: 81
i: [1,50,0,5,1,0,0,1] t4_boole
s: 50
i: [1,51,0,5,1,0,0,1] t3_boole
s: 51
i: [1,52,0,5,2,0,0,2] t6_boole
s: 52
i: [1,57,0,5,2,0,0,1] t7_boole
s: 57
d: [2,99,0,5,1,1,22,50,86]
s: 99
d: [2,100,0,5,1,1,22,51,86]
s: 100
d: [2,101,0,5,2,1,22,52,86]
s: 101
i: [1,59,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 59
i: [1,60,0,6,2,0,0,2] symmetry_r1_xboole_0
s: 60
i: [1,53,0,7,2,0,0,1] l20_zfmisc_1
s: 53
i: [1,56,0,7,3,0,0,1] t8_boole
s: 56
i: [1,58,0,7,2,0,0,1] l18_zfmisc_1
s: 58
i: [1,54,0,8,2,0,0,1] t83_xboole_1
s: 54
d: [2,108,1,9,2,1,44,66,54]
d: [2,109,1,4,1,1,19,108]
a: [3,113,1,4,1,109]
s: 113
d: [2,115,2,0,0,2,52,69,113,53]
e: 115
d: [2,116,2,3,1,1,40,113,53]
d: [2,118,2,0,0,2,21,116,69]
e: 118
r: 66
r: 69
a: [3,70,0,3,1,45]
s: 70
d: [2,122,1,2,1,1,40,70,57]
a: [3,131,1,2,1,122]
s: 131
a: [3,65,0,7,1,45]
s: 65
i: [1,55,0,8,2,0,0,1] t83_xboole_1
s: 55
d: [2,133,2,0,0,2,52,113,65,55]
e: 133
d: [2,137,1,9,2,1,44,55,65]
d: [2,138,1,4,1,1,19,137]
d: [2,142,1,0,0,2,21,138,113]
e: 142
r: 113
d: [2,121,1,3,1,1,40,70,59]
a: [3,126,1,3,1,121]
s: 126
a: [3,112,1,4,1,109]
s: 112
d: [2,145,2,0,0,2,52,70,112,58]
e: 145
d: [2,146,2,3,1,1,40,112,58]
d: [2,149,2,0,0,2,21,146,70]
e: 149
i: [1,44,0,10,2,1,0,0] t59_zfmisc_1
f: 149,138,116,109,44,45
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t59_zfmisc_1
% SZS output start Proof for zfmisc_1__t59_zfmisc_1
1. ! [X0] : ! [X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) <=> ~r2_hidden(X0,X1)) [input t59_zfmisc_1]
2. ~! [X0] : ! [X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) <=> ~r2_hidden(X0,X1)) [negated conjecture 1]
19. ! [X0] : ! [X1] : ~(r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [input l18_zfmisc_1]
20. ! [X0] : ! [X1] : (~r2_hidden(X0,X1) => r1_xboole_0(k1_tarski(X0),X1)) [input l20_zfmisc_1]
21. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [input t83_xboole_1]
22. ~! [X0,X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) <=> ~r2_hidden(X0,X1)) [flattening 2]
23. ! [X0,X1] : (~r2_hidden(X0,X1) => r1_xboole_0(k1_tarski(X0),X1)) [flattening 20]
24. ! [X0,X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [flattening 21]
27. ! [X0,X1] : ~(r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [flattening 19]
28. ? [X0,X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) <~> ~r2_hidden(X0,X1)) [ennf transformation 22]
30. ! [X0,X1] : (r1_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1)) [ennf transformation 23]
33. ! [X0,X1] : (~r2_hidden(X0,X1) | ~r1_xboole_0(k1_tarski(X0),X1)) [ennf transformation 27]
36. ? [X0,X1] : ((r2_hidden(X0,X1) | k1_tarski(X0) != k4_xboole_0(k1_tarski(X0),X1)) & (~r2_hidden(X0,X1) | k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1))) [nnf transformation 28]
37. ? [X0,X1] : ((r2_hidden(X0,X1) | k1_tarski(X0) != k4_xboole_0(k1_tarski(X0),X1)) & (~r2_hidden(X0,X1) | k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1))) => ((r2_hidden(sK0,sK1) | k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),sK1)) & (~r2_hidden(sK0,sK1) | k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),sK1))) [choice axiom]
38. (r2_hidden(sK0,sK1) | k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),sK1)) & (~r2_hidden(sK0,sK1) | k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),sK1)) [skolemisation 36,37]
39. ! [X0,X1] : ((r1_xboole_0(X0,X1) | k4_xboole_0(X0,X1) != X0) & (k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1))) [nnf transformation 24]
44. ~r2_hidden(sK0,sK1) | k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),sK1) [cnf transformation 38] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. r2_hidden(sK0,sK1) | k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),sK1) [cnf transformation 38] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. r1_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1) [cnf transformation 30] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1) [cnf transformation 39] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. k4_xboole_0(X0,X1) != X0 | r1_xboole_0(X0,X1) [cnf transformation 39] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. ~r1_xboole_0(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 33] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. 1 <=> k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),sK1) [avatar definition]
65. k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),sK1) <- (1) [avatar component clause 64] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66. k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),sK1) <- (~1) [avatar component clause 64] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. 2 <=> r2_hidden(sK0,sK1) [avatar definition]
70. r2_hidden(sK0,sK1) <- (2) [avatar component clause 68] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. ~1 | 2 [avatar split clause 45,68,64]
72. 1 | ~2 [avatar split clause 44,68,64]
108. k1_tarski(sK0) != k1_tarski(sK0) | ~r1_xboole_0(k1_tarski(sK0),sK1) <- (~1) [superposition 66,54] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
109. ~r1_xboole_0(k1_tarski(sK0),sK1) <- (~1) [trivial inequality removal 108] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
111. 8 <=> r1_xboole_0(k1_tarski(sK0),sK1) [avatar definition]
112. r1_xboole_0(k1_tarski(sK0),sK1) <- (8) [avatar component clause 111] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. ~r1_xboole_0(k1_tarski(sK0),sK1) <- (~8) [avatar component clause 111] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. ~8 | 1 [avatar split clause 109,64,111]
116. r2_hidden(sK0,sK1) <- (~8) [resolution 113,53] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
120. 2 | 8 [avatar split clause 116,111,68]
137. k1_tarski(sK0) != k1_tarski(sK0) | r1_xboole_0(k1_tarski(sK0),sK1) <- (1) [superposition 55,65] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
138. r1_xboole_0(k1_tarski(sK0),sK1) <- (1) [trivial inequality removal 137] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
144. 8 | ~1 [avatar split clause 138,64,111]
146. ~r2_hidden(sK0,sK1) <- (8) [resolution 112,58] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
149. $false <- (2, 8) [subsumption resolution 146,70] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
150. ~2 | ~8 [avatar contradiction clause 149]
151. $false [avatar sat refutation 71,72,114,120,144,150] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:75,thDist:-75}
% SZS output end Proof for zfmisc_1__t59_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 66
% Active clauses: 25
% Passive clauses: 34
% Final active clauses: 21
% Final passive clauses: 6
% Main loop iterations started: 28
% 
% >>> Simplifying Inferences
% Trivial inequalities: 4
% Fw subsumption resolutions: 3
% Fw demodulations: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 7
% Unit resulting resolution: 3
% Forward superposition: 4
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 5
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

