i: [1,124,0,2,1,1,0,0] t18_compts_1
a: [3,176,0,2,1,124]
s: 176
i: [1,122,0,2,1,1,0,0] t18_compts_1
a: [3,186,0,2,1,122]
s: 186
i: [1,121,0,2,1,1,0,0] t18_compts_1
a: [3,191,0,2,1,121]
s: 191
i: [1,123,0,3,1,1,0,0] t18_compts_1
a: [3,181,0,3,1,123]
s: 181
i: [1,125,0,2,1,0,0,255] spc1_boole
a: [3,196,0,2,1,125]
s: 196
i: [1,126,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,201,0,2,1,126]
s: 201
i: [1,127,0,2,1,0,0,255] fc1_xboole_0
a: [3,206,0,2,1,127]
s: 206
i: [1,164,0,2,1,0,0,2] rc2_xboole_0
a: [3,216,0,2,1,164]
s: 216
i: [1,165,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,221,0,2,1,165]
s: 221
i: [1,166,0,2,1,0,0,2] existence_l1_struct_0
a: [3,226,0,2,1,166]
s: 226
i: [1,167,0,2,1,0,0,2] rc1_xboole_0
a: [3,231,0,2,1,167]
s: 231
i: [1,169,0,2,1,0,0,1] rc1_finset_1
a: [3,236,0,2,1,169]
s: 236
i: [1,168,0,2,1,0,0,1] rc1_finset_1
a: [3,241,0,2,1,168]
s: 241
i: [1,172,0,2,1,0,0,255] rc8_finset_1
a: [3,246,0,2,1,172]
s: 246
i: [1,171,0,2,1,0,0,255] rc8_finset_1
a: [3,251,0,2,1,171]
s: 251
i: [1,170,0,2,1,0,0,255] rc8_finset_1
a: [3,256,0,2,1,170]
s: 256
i: [1,128,0,3,1,0,0,255] d2_xboole_0
a: [3,211,0,3,1,128]
s: 211
r: 201
i: [1,129,0,4,2,0,0,255] cc5_finset_1
s: 129
i: [1,130,0,4,2,0,0,255] cc4_finset_1
s: 130
i: [1,132,0,4,2,0,0,1] dt_l1_pre_topc
s: 132
i: [1,149,0,4,2,0,0,255] cc6_finset_1
s: 149
d: [2,259,1,2,1,1,40,132,186]
a: [3,264,1,2,1,259]
s: 264
i: [1,150,0,4,2,0,0,1] cc1_finset_1
s: 150
d: [2,260,1,2,1,1,40,132,221]
a: [3,269,1,2,1,260]
s: 269
i: [1,162,0,4,1,0,0,255] existence_m1_subset_1
s: 162
i: [1,151,0,5,2,0,0,255] t6_boole
s: 151
d: [2,271,1,2,1,1,40,149,231]
a: [3,276,1,2,1,271]
s: 276
i: [1,161,0,5,2,0,0,255] t7_boole
s: 161
d: [2,272,1,2,1,1,40,149,206]
a: [3,281,1,2,1,272]
s: 281
i: [1,135,0,6,3,0,0,2] cc7_struct_0
s: 135
i: [1,136,0,6,3,0,0,2] cc5_struct_0
s: 136
d: [2,283,1,2,1,1,40,150,231]
a: [3,288,1,2,1,283]
s: 288
i: [1,138,0,6,3,0,0,3] cc2_struct_0
s: 138
d: [2,284,1,2,1,1,40,150,206]
a: [3,293,1,2,1,284]
s: 293
i: [1,139,0,6,3,0,0,3] cc17_struct_0
s: 139
i: [1,140,0,6,3,0,0,3] cc1_struct_0
s: 140
i: [1,143,0,6,3,0,0,2] cc4_struct_0
s: 143
i: [1,144,0,6,3,0,0,2] cc6_struct_0
s: 144
i: [1,158,0,6,2,0,0,255] t1_subset
s: 158
i: [1,163,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 163
i: [1,131,0,7,3,0,0,255] cc7_finset_1
s: 131
d: [2,299,1,5,2,0,40,131,162]
s: 299
i: [1,141,0,7,3,0,0,255] cc8_struct_0
s: 141
i: [1,145,0,7,3,0,0,255] cc11_struct_0
s: 145
i: [1,146,0,7,3,0,0,255] cc11_struct_0
s: 146
i: [1,147,0,7,3,0,0,255] cc9_struct_0
s: 147
i: [1,152,0,7,3,0,0,1] fc2_struct_0
s: 152
i: [1,153,0,7,3,0,0,1] fc9_struct_0
s: 153
d: [2,304,1,4,2,1,40,153,181]
d: [2,305,1,2,1,2,21,304,264]
a: [3,309,1,2,1,305]
s: 309
i: [1,154,0,7,3,0,0,255] fc6_struct_0
s: 154
i: [1,155,0,7,3,0,0,1] fc1_struct_0
s: 155
i: [1,156,0,7,3,0,0,1] fc8_struct_0
s: 156
i: [1,157,0,7,3,0,0,255] fc7_struct_0
s: 157
i: [1,160,0,7,3,0,0,255] t8_boole
s: 160
i: [1,134,0,8,4,0,0,1] cc3_compts_1
s: 134
d: [2,320,2,0,0,4,52,186,176,191,309,134]
e: 320
d: [2,321,2,6,3,1,40,134,309]
d: [2,323,2,4,2,2,21,321,191]
d: [2,324,2,2,1,3,21,323,176]
d: [2,325,2,0,0,4,21,324,186]
e: 325
f: 325,305,259,121,122,123,124
% Refutation found. Thanks to Tanya!
% SZS status Theorem for compts_1__t18_compts_1
% SZS output start Proof for compts_1__t18_compts_1
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => (v1_finset_1(u1_struct_0(X0)) => v1_compts_1(X0))) [input t18_compts_1]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => (v1_finset_1(u1_struct_0(X0)) => v1_compts_1(X0))) [negated conjecture 1]
10. ! [X0] : (l1_pre_topc(X0) => ((v2_pre_topc(X0) & v8_struct_0(X0)) => (v1_compts_1(X0) & v2_pre_topc(X0)))) [input cc3_compts_1]
24. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
38. ! [X0] : ((l1_struct_0(X0) & ~v8_struct_0(X0)) => ~v1_finset_1(u1_struct_0(X0))) [input fc9_struct_0]
54. ? [X0] : ((~v1_compts_1(X0) & v1_finset_1(u1_struct_0(X0))) & (l1_pre_topc(X0) & v2_pre_topc(X0))) [ennf transformation 2]
55. ? [X0] : (~v1_compts_1(X0) & v1_finset_1(u1_struct_0(X0)) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 54]
59. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 24]
60. ! [X0] : (((v1_compts_1(X0) & v2_pre_topc(X0)) | (~v2_pre_topc(X0) | ~v8_struct_0(X0))) | ~l1_pre_topc(X0)) [ennf transformation 10]
61. ! [X0] : ((v1_compts_1(X0) & v2_pre_topc(X0)) | ~v2_pre_topc(X0) | ~v8_struct_0(X0) | ~l1_pre_topc(X0)) [flattening 60]
89. ! [X0] : (~v1_finset_1(u1_struct_0(X0)) | (~l1_struct_0(X0) | v8_struct_0(X0))) [ennf transformation 38]
90. ! [X0] : (~v1_finset_1(u1_struct_0(X0)) | ~l1_struct_0(X0) | v8_struct_0(X0)) [flattening 89]
105. ? [X0] : (~v1_compts_1(X0) & v1_finset_1(u1_struct_0(X0)) & l1_pre_topc(X0) & v2_pre_topc(X0)) => (~v1_compts_1(sK0) & v1_finset_1(u1_struct_0(sK0)) & l1_pre_topc(sK0) & v2_pre_topc(sK0)) [choice axiom]
106. ~v1_compts_1(sK0) & v1_finset_1(u1_struct_0(sK0)) & l1_pre_topc(sK0) & v2_pre_topc(sK0) [skolemisation 55,105]
121. v2_pre_topc(sK0) [cnf transformation 106] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
122. l1_pre_topc(sK0) [cnf transformation 106] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
123. v1_finset_1(u1_struct_0(sK0)) [cnf transformation 106] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
124. ~v1_compts_1(sK0) [cnf transformation 106] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
132. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 59] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
134. ~v8_struct_0(X0) | ~v2_pre_topc(X0) | v1_compts_1(X0) | ~l1_pre_topc(X0) [cnf transformation 61] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
153. ~v1_finset_1(u1_struct_0(X0)) | ~l1_struct_0(X0) | v8_struct_0(X0) [cnf transformation 90] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
174. 1 <=> v1_compts_1(sK0) [avatar definition]
176. ~v1_compts_1(sK0) <- (~1) [avatar component clause 174] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
177. ~1 [avatar split clause 124,174]
179. 2 <=> v1_finset_1(u1_struct_0(sK0)) [avatar definition]
181. v1_finset_1(u1_struct_0(sK0)) <- (2) [avatar component clause 179] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
182. 2 [avatar split clause 123,179]
184. 3 <=> l1_pre_topc(sK0) [avatar definition]
186. l1_pre_topc(sK0) <- (3) [avatar component clause 184] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
187. 3 [avatar split clause 122,184]
189. 4 <=> v2_pre_topc(sK0) [avatar definition]
191. v2_pre_topc(sK0) <- (4) [avatar component clause 189] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
192. 4 [avatar split clause 121,189]
259. l1_struct_0(sK0) <- (3) [resolution 132,186] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
262. 18 <=> l1_struct_0(sK0) [avatar definition]
264. l1_struct_0(sK0) <- (18) [avatar component clause 262] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
265. 18 | ~3 [avatar split clause 259,184,262]
304. ~l1_struct_0(sK0) | v8_struct_0(sK0) <- (2) [resolution 153,181] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
305. v8_struct_0(sK0) <- (2, 18) [subsumption resolution 304,264] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
307. 24 <=> v8_struct_0(sK0) [avatar definition]
309. v8_struct_0(sK0) <- (24) [avatar component clause 307] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
310. 24 | ~2 | ~18 [avatar split clause 305,262,179,307]
321. ~v2_pre_topc(sK0) | v1_compts_1(sK0) | ~l1_pre_topc(sK0) <- (24) [resolution 134,309] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
323. v1_compts_1(sK0) | ~l1_pre_topc(sK0) <- (4, 24) [subsumption resolution 321,191] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
324. ~l1_pre_topc(sK0) <- (~1, 4, 24) [subsumption resolution 323,176] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
325. $false <- (~1, 3, 4, 24) [subsumption resolution 324,186] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
326. 1 | ~3 | ~4 | ~24 [avatar contradiction clause 325]
327. $false [avatar sat refutation 177,182,187,192,265,310,326] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:95,thDist:-95}
% SZS output end Proof for compts_1__t18_compts_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 47
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 52
% Generated clauses: 108
% Active clauses: 50
% Passive clauses: 63
% Final active clauses: 49
% Final passive clauses: 8
% Main loop iterations started: 55
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 16
% Unit resulting resolution: 1
% Forward superposition: 4
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 28
% SAT solver unit clauses: 17
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

