i: [1,92,0,2,1,1,0,0] t12_setwiseo
a: [3,136,0,2,1,92]
s: 136
i: [1,94,0,3,1,1,0,0] t12_setwiseo
a: [3,126,0,3,1,94]
s: 126
i: [1,93,0,4,1,1,0,0] t12_setwiseo
a: [3,131,0,4,1,93]
s: 131
i: [1,95,0,6,2,1,0,0] t12_setwiseo
s: 95
i: [1,96,0,2,1,0,0,1] fc1_xboole_0
a: [3,141,0,2,1,96]
s: 141
i: [1,116,0,2,1,0,0,1] rc2_xboole_0
a: [3,146,0,2,1,116]
s: 146
i: [1,117,0,2,1,0,0,1] rc1_xboole_0
a: [3,151,0,2,1,117]
s: 151
i: [1,119,0,2,1,0,0,255] rc1_relat_1
a: [3,156,0,2,1,119]
s: 156
i: [1,118,0,2,1,0,0,255] rc1_relat_1
a: [3,161,0,2,1,118]
s: 161
i: [1,120,0,2,1,0,0,2] rc1_finset_1
a: [3,166,0,2,1,120]
s: 166
i: [1,121,0,2,1,0,0,255] rc8_finset_1
a: [3,171,0,2,1,121]
s: 171
i: [1,122,0,2,1,0,0,255] rc2_relat_1
a: [3,176,0,2,1,122]
s: 176
i: [1,97,0,3,1,0,0,255] dt_k5_finsub_1
s: 97
i: [1,98,0,3,1,0,0,255] fc2_finsub_1
s: 98
i: [1,99,0,3,1,0,0,255] fc2_finsub_1
s: 99
i: [1,101,0,4,2,0,0,255] cc1_finsub_1
s: 101
d: [2,178,1,3,1,0,40,101,97]
s: 178
i: [1,102,0,4,2,0,0,255] cc1_finsub_1
s: 102
d: [2,179,1,3,1,0,40,102,97]
s: 179
i: [1,103,0,4,2,0,0,255] cc1_relat_1
s: 103
i: [1,113,0,4,1,0,0,1] existence_m1_subset_1
s: 113
d: [2,180,1,2,1,1,40,103,141]
a: [3,185,1,2,1,180]
s: 185
i: [1,104,0,5,2,0,0,1] t6_boole
s: 104
d: [2,181,1,2,1,1,40,103,151]
a: [3,190,1,2,1,181]
s: 190
i: [1,112,0,5,2,0,0,1] t7_boole
s: 112
i: [1,105,0,6,3,0,0,255] cc2_finsub_1
s: 105
i: [1,109,0,6,2,0,0,1] t1_subset
s: 109
i: [1,114,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 114
i: [1,111,0,7,3,0,0,1] t8_boole
s: 111
i: [1,110,0,8,3,0,0,1] t2_subset
s: 110
i: [1,115,0,11,3,0,0,255] dt_o_2_0_setwiseo
s: 115
d: [2,196,1,7,2,1,40,110,131]
d: [2,198,1,4,1,1,21,196,98]
a: [3,202,1,4,1,198]
s: 202
i: [1,100,0,12,4,0,0,1] t9_setwiseo
s: 100
d: [2,213,1,8,3,1,40,100,131]
d: [2,216,1,6,2,2,21,213,136]
d: [2,217,1,3,1,2,21,216,95]
s: 217
r: 95
d: [2,197,1,6,2,0,40,110,113]
s: 197
d: [2,192,1,7,3,0,44,104,104]
s: 192
d: [2,221,2,2,1,2,40,197,217]
a: [3,225,2,2,1,221]
s: 225
d: [2,227,3,0,0,3,52,141,126,225,111]
e: 227
d: [2,228,3,0,0,3,52,141,126,225,111]
e: 228
d: [2,229,3,0,0,2,52,126,225,104]
e: 229
f: 229,221,92,93,94
% Refutation found. Thanks to Tanya!
% SZS status Theorem for setwiseo__t12_setwiseo
% SZS output start Proof for setwiseo__t12_setwiseo
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,k5_finsub_1(X0)) => ~(! [X2] : (m1_subset_1(X2,X0) => ~r2_hidden(X2,X1)) & k1_xboole_0 != X1))) [input t12_setwiseo]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,k5_finsub_1(X0)) => ~(! [X2] : (m1_subset_1(X2,X0) => ~r2_hidden(X2,X1)) & k1_xboole_0 != X1))) [negated conjecture 1]
20. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
30. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
31. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
34. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,k5_finsub_1(X0)) => ! [X2] : (r2_hidden(X2,X1) => m1_subset_1(X2,X0)))) [input t9_setwiseo]
36. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 30]
53. ? [X0] : (? [X1] : ((! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,X0)) & k1_xboole_0 != X1) & m1_subset_1(X1,k5_finsub_1(X0))) & ~v1_xboole_0(X0)) [ennf transformation 2]
54. ? [X0] : (? [X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,X0)) & k1_xboole_0 != X1 & m1_subset_1(X1,k5_finsub_1(X0))) & ~v1_xboole_0(X0)) [flattening 53]
55. ! [X0] : (! [X1] : (! [X2] : (m1_subset_1(X2,X0) | ~r2_hidden(X2,X1)) | ~m1_subset_1(X1,k5_finsub_1(X0))) | v1_xboole_0(X0)) [ennf transformation 34]
58. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 31]
68. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 36]
69. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 68]
75. ? [X0] : (? [X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,X0)) & k1_xboole_0 != X1 & m1_subset_1(X1,k5_finsub_1(X0))) & ~v1_xboole_0(X0)) => (? [X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,sK0)) & k1_xboole_0 != X1 & m1_subset_1(X1,k5_finsub_1(sK0))) & ~v1_xboole_0(sK0)) [choice axiom]
76. ? [X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,sK0)) & k1_xboole_0 != X1 & m1_subset_1(X1,k5_finsub_1(sK0))) => (! [X2] : (~r2_hidden(X2,sK1) | ~m1_subset_1(X2,sK0)) & k1_xboole_0 != sK1 & m1_subset_1(sK1,k5_finsub_1(sK0))) [choice axiom]
77. (! [X2] : (~r2_hidden(X2,sK1) | ~m1_subset_1(X2,sK0)) & k1_xboole_0 != sK1 & m1_subset_1(sK1,k5_finsub_1(sK0))) & ~v1_xboole_0(sK0) [skolemisation 54,76,75]
78. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK2(X0),X0)) [choice axiom]
79. ! [X0] : m1_subset_1(sK2(X0),X0) [skolemisation 20,78]
92. ~v1_xboole_0(sK0) [cnf transformation 77] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
93. m1_subset_1(sK1,k5_finsub_1(sK0)) [cnf transformation 77] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
94. k1_xboole_0 != sK1 [cnf transformation 77] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
95. ~r2_hidden(X2,sK1) | ~m1_subset_1(X2,sK0) [cnf transformation 77] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. ~m1_subset_1(X1,k5_finsub_1(X0)) | ~r2_hidden(X2,X1) | m1_subset_1(X2,X0) | v1_xboole_0(X0) [cnf transformation 55] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 58] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
110. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 69] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
113. m1_subset_1(sK2(X0),X0) [cnf transformation 79] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
124. 1 <=> k1_xboole_0 = sK1 [avatar definition]
126. k1_xboole_0 != sK1 <- (~1) [avatar component clause 124] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
127. ~1 [avatar split clause 94,124]
129. 2 <=> m1_subset_1(sK1,k5_finsub_1(sK0)) [avatar definition]
131. m1_subset_1(sK1,k5_finsub_1(sK0)) <- (2) [avatar component clause 129] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
132. 2 [avatar split clause 93,129]
134. 3 <=> v1_xboole_0(sK0) [avatar definition]
136. ~v1_xboole_0(sK0) <- (~3) [avatar component clause 134] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
137. ~3 [avatar split clause 92,134]
197. r2_hidden(sK2(X0),X0) | v1_xboole_0(X0) [resolution 110,113] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
213. ~r2_hidden(X0,sK1) | m1_subset_1(X0,sK0) | v1_xboole_0(sK0) <- (2) [resolution 100,131] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
216. ~r2_hidden(X0,sK1) | m1_subset_1(X0,sK0) <- (2, ~3) [subsumption resolution 213,136] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
217. ~r2_hidden(X0,sK1) <- (2, ~3) [subsumption resolution 216,95] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
221. v1_xboole_0(sK1) <- (2, ~3) [resolution 197,217] {a:2,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
223. 16 <=> v1_xboole_0(sK1) [avatar definition]
225. v1_xboole_0(sK1) <- (16) [avatar component clause 223] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
226. 16 | ~2 | 3 [avatar split clause 221,134,129,223]
229. $false <- (~1, 16) [unit resulting resolution 126,225,104] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
233. 1 | ~16 [avatar contradiction clause 229]
234. $false [avatar sat refutation 127,132,137,226,233] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:102,thDist:-102}
% SZS output end Proof for setwiseo__t12_setwiseo
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 33
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 6
% 
% >>> Saturation
% Initial clauses: 31
% Generated clauses: 75
% Active clauses: 35
% Passive clauses: 42
% Final active clauses: 34
% Final passive clauses: 5
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 7
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 19
% Unit resulting resolution: 3
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 11
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

