i: [1,114,0,2,1,1,0,0] t25_yellow_2
a: [3,159,0,2,1,114]
s: 159
i: [1,112,0,2,1,1,0,0] t25_yellow_2
a: [3,164,0,2,1,112]
s: 164
i: [1,111,0,2,1,1,0,0] t25_yellow_2
a: [3,169,0,2,1,111]
s: 169
i: [1,113,0,3,1,1,0,0] t25_yellow_2
s: 113
i: [1,115,0,2,1,0,0,255] spc1_boole
a: [3,174,0,2,1,115]
s: 174
i: [1,116,0,2,1,0,0,255] fc1_xboole_0
a: [3,179,0,2,1,116]
s: 179
i: [1,152,0,2,1,0,0,255] rc2_xboole_0
a: [3,184,0,2,1,152]
s: 184
i: [1,153,0,2,1,0,0,1] existence_l1_orders_2
a: [3,189,0,2,1,153]
s: 189
i: [1,154,0,2,1,0,0,2] existence_l1_struct_0
a: [3,194,0,2,1,154]
s: 194
i: [1,155,0,2,1,0,0,255] rc1_xboole_0
a: [3,199,0,2,1,155]
s: 199
i: [1,117,0,4,2,0,0,1] dt_l1_orders_2
s: 117
d: [2,201,1,2,1,1,40,117,164]
a: [3,206,1,2,1,201]
s: 206
i: [1,150,0,4,1,0,0,255] existence_m1_subset_1
s: 150
d: [2,202,1,2,1,1,40,117,189]
a: [3,211,1,2,1,202]
s: 211
i: [1,142,0,5,2,0,0,255] t6_boole
s: 142
i: [1,149,0,5,2,0,0,255] t7_boole
s: 149
i: [1,118,0,6,3,0,0,255] cc2_lattice3
s: 118
i: [1,119,0,6,3,0,0,255] cc1_lattice3
s: 119
i: [1,120,0,6,3,0,0,255] cc4_yellow_0
s: 120
i: [1,121,0,6,3,0,0,255] cc4_yellow_0
s: 121
i: [1,128,0,6,3,0,0,255] cc7_struct_0
s: 128
i: [1,129,0,6,3,0,0,255] cc5_struct_0
s: 129
i: [1,131,0,6,3,0,0,255] cc17_struct_0
s: 131
d: [2,215,1,4,2,1,40,129,194]
a: [3,221,1,2,1,215]
s: 221
i: [1,132,0,6,3,0,0,255] cc2_struct_0
s: 132
i: [1,133,0,6,3,0,0,255] cc6_struct_0
s: 133
d: [2,217,2,4,2,1,40,129,211]
a: [3,230,2,2,1,217]
s: 230
i: [1,134,0,6,3,0,0,255] cc1_struct_0
s: 134
i: [1,137,0,6,3,0,0,255] cc4_struct_0
s: 137
i: [1,146,0,6,2,0,0,255] t1_subset
s: 146
i: [1,151,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 151
i: [1,135,0,7,3,0,0,255] cc8_struct_0
s: 135
i: [1,138,0,7,3,0,0,255] cc11_struct_0
s: 138
i: [1,139,0,7,3,0,0,255] cc11_struct_0
s: 139
i: [1,140,0,7,3,0,0,255] cc9_struct_0
s: 140
i: [1,148,0,7,3,0,0,255] t8_boole
s: 148
d: [2,213,1,7,3,0,44,142,142]
s: 213
i: [1,123,0,8,4,0,0,255] cc1_yellow_0
s: 123
i: [1,124,0,8,4,0,0,255] cc1_yellow_0
s: 124
i: [1,126,0,8,4,0,0,255] cc3_yellow_0
s: 126
i: [1,127,0,8,4,0,0,255] cc5_yellow_0
s: 127
i: [1,147,0,8,3,0,0,255] t2_subset
s: 147
i: [1,141,0,9,4,0,0,255] cc10_struct_0
s: 141
d: [2,253,1,6,2,0,40,147,150]
s: 253
i: [1,143,0,10,4,0,0,2] t24_yellow_2
s: 143
i: [1,144,0,11,4,0,0,1] t23_yellow_2
s: 144
d: [2,258,2,6,2,0,40,253,151]
s: 258
i: [1,145,0,11,4,0,0,1] t23_yellow_2
s: 145
d: [2,261,1,7,3,0,40,145,113]
d: [2,262,1,5,2,1,21,261,164]
d: [2,263,1,3,1,2,21,262,169]
s: 263
d: [2,264,2,0,0,3,52,169,164,159,263,143]
e: 264
d: [2,265,2,6,3,2,40,263,143]
d: [2,268,2,4,2,3,21,265,159]
d: [2,269,2,2,1,3,21,268,164]
d: [2,270,2,0,0,3,21,269,169]
e: 270
f: 270,111,112,114
% Refutation found. Thanks to Tanya!
% SZS status Theorem for yellow_2__t25_yellow_2
% SZS output start Proof for yellow_2__t25_yellow_2
1. ! [X0] : ((l1_orders_2(X0) & ~v2_struct_0(X0)) => (! [X1] : r2_yellow_0(X0,X1) => v3_lattice3(X0))) [input t25_yellow_2]
2. ~! [X0] : ((l1_orders_2(X0) & ~v2_struct_0(X0)) => (! [X1] : r2_yellow_0(X0,X1) => v3_lattice3(X0))) [negated conjecture 1]
33. ! [X0] : ((l1_orders_2(X0) & ~v2_struct_0(X0)) => (! [X1] : r1_yellow_0(X0,X1) <=> ! [X1] : r2_yellow_0(X0,X1))) [input t23_yellow_2]
34. ! [X0] : ((l1_orders_2(X0) & ~v2_struct_0(X0)) => (! [X1] : r1_yellow_0(X0,X1) => v3_lattice3(X0))) [input t24_yellow_2]
39. ! [X0] : ((l1_orders_2(X0) & ~v2_struct_0(X0)) => (! [X1] : r1_yellow_0(X0,X1) <=> ! [X2] : r2_yellow_0(X0,X2))) [rectify 33]
44. ? [X0] : ((~v3_lattice3(X0) & ! [X1] : r2_yellow_0(X0,X1)) & (l1_orders_2(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
45. ? [X0] : (~v3_lattice3(X0) & ! [X1] : r2_yellow_0(X0,X1) & l1_orders_2(X0) & ~v2_struct_0(X0)) [flattening 44]
82. ! [X0] : ((v3_lattice3(X0) | ? [X1] : ~r1_yellow_0(X0,X1)) | (~l1_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 34]
83. ! [X0] : (v3_lattice3(X0) | ? [X1] : ~r1_yellow_0(X0,X1) | ~l1_orders_2(X0) | v2_struct_0(X0)) [flattening 82]
84. ! [X0] : ((! [X1] : r1_yellow_0(X0,X1) <=> ! [X2] : r2_yellow_0(X0,X2)) | (~l1_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 39]
85. ! [X0] : ((! [X1] : r1_yellow_0(X0,X1) <=> ! [X2] : r2_yellow_0(X0,X2)) | ~l1_orders_2(X0) | v2_struct_0(X0)) [flattening 84]
92. ? [X0] : (~v3_lattice3(X0) & ! [X1] : r2_yellow_0(X0,X1) & l1_orders_2(X0) & ~v2_struct_0(X0)) => (~v3_lattice3(sK0) & ! [X1] : r2_yellow_0(sK0,X1) & l1_orders_2(sK0) & ~v2_struct_0(sK0)) [choice axiom]
93. ~v3_lattice3(sK0) & ! [X1] : r2_yellow_0(sK0,X1) & l1_orders_2(sK0) & ~v2_struct_0(sK0) [skolemisation 45,92]
94. ! [X0] : (? [X1] : ~r1_yellow_0(X0,X1) => ~r1_yellow_0(X0,sK1(X0))) [choice axiom]
95. ! [X0] : (v3_lattice3(X0) | ~r1_yellow_0(X0,sK1(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0)) [skolemisation 83,94]
96. ! [X0] : (((! [X1] : r1_yellow_0(X0,X1) | ? [X2] : ~r2_yellow_0(X0,X2)) & (! [X2] : r2_yellow_0(X0,X2) | ? [X1] : ~r1_yellow_0(X0,X1))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [nnf transformation 85]
97. ! [X0] : (((! [X1] : r1_yellow_0(X0,X1) | ? [X2] : ~r2_yellow_0(X0,X2)) & (! [X3] : r2_yellow_0(X0,X3) | ? [X4] : ~r1_yellow_0(X0,X4))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [rectify 96]
98. ! [X0] : (? [X2] : ~r2_yellow_0(X0,X2) => ~r2_yellow_0(X0,sK2(X0))) [choice axiom]
99. ! [X0] : (? [X4] : ~r1_yellow_0(X0,X4) => ~r1_yellow_0(X0,sK3(X0))) [choice axiom]
100. ! [X0] : (((! [X1] : r1_yellow_0(X0,X1) | ~r2_yellow_0(X0,sK2(X0))) & (! [X3] : r2_yellow_0(X0,X3) | ~r1_yellow_0(X0,sK3(X0)))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [skolemisation 97,99,98]
111. ~v2_struct_0(sK0) [cnf transformation 93] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
112. l1_orders_2(sK0) [cnf transformation 93] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
113. r2_yellow_0(sK0,X1) [cnf transformation 93] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
114. ~v3_lattice3(sK0) [cnf transformation 93] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
143. ~r1_yellow_0(X0,sK1(X0)) | v3_lattice3(X0) | ~l1_orders_2(X0) | v2_struct_0(X0) [cnf transformation 95] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
145. ~r2_yellow_0(X0,sK2(X0)) | r1_yellow_0(X0,X1) | ~l1_orders_2(X0) | v2_struct_0(X0) [cnf transformation 100] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
157. 1 <=> v3_lattice3(sK0) [avatar definition]
159. ~v3_lattice3(sK0) <- (~1) [avatar component clause 157] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
160. ~1 [avatar split clause 114,157]
162. 2 <=> l1_orders_2(sK0) [avatar definition]
164. l1_orders_2(sK0) <- (2) [avatar component clause 162] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
165. 2 [avatar split clause 112,162]
167. 3 <=> v2_struct_0(sK0) [avatar definition]
169. ~v2_struct_0(sK0) <- (~3) [avatar component clause 167] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
170. ~3 [avatar split clause 111,167]
261. r1_yellow_0(sK0,X0) | ~l1_orders_2(sK0) | v2_struct_0(sK0) [resolution 145,113] {a:1,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
262. r1_yellow_0(sK0,X0) | v2_struct_0(sK0) <- (2) [subsumption resolution 261,164] {a:1,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
263. r1_yellow_0(sK0,X0) <- (2, ~3) [subsumption resolution 262,169] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
265. v3_lattice3(sK0) | ~l1_orders_2(sK0) | v2_struct_0(sK0) <- (2, ~3) [resolution 263,143] {a:2,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
268. ~l1_orders_2(sK0) | v2_struct_0(sK0) <- (~1, 2, ~3) [subsumption resolution 265,159] {a:2,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
269. v2_struct_0(sK0) <- (~1, 2, ~3) [subsumption resolution 268,164] {a:2,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
270. $false <- (~1, 2, ~3) [subsumption resolution 269,169] {a:2,w:0,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
271. 1 | ~2 | 3 [avatar contradiction clause 270]
272. $false [avatar sat refutation 160,165,170,271] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:81,thDist:-81}
% SZS output end Proof for yellow_2__t25_yellow_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 37
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 45
% Generated clauses: 89
% Active clauses: 44
% Passive clauses: 51
% Final active clauses: 44
% Final passive clauses: 2
% Main loop iterations started: 49
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 18
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 19
% SAT solver unit clauses: 9
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

