i: [1,37,0,2,1,1,0,0] t64_orders_1
a: [3,74,0,2,1,37]
s: 74
i: [1,41,0,3,1,1,0,0] t64_orders_1
s: 41
i: [1,40,0,3,1,1,0,0] t64_orders_1
a: [3,59,0,3,1,40]
s: 59
i: [1,38,0,3,1,1,0,0] t64_orders_1
a: [3,69,0,3,1,38]
s: 69
i: [1,39,0,4,1,1,0,0] t64_orders_1
a: [3,64,0,4,1,39]
s: 64
i: [1,42,0,5,2,0,0,2] dt_k2_relat_1
s: 42
i: [1,52,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 52
i: [1,43,0,7,2,0,0,2] involutiveness_k2_relat_1
s: 43
i: [1,53,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 53
i: [1,44,0,8,2,0,0,1] t21_relat_1
s: 44
i: [1,54,0,8,3,0,0,255] fc3_relat_1
s: 54
i: [1,46,0,9,3,0,0,1] t41_orders_1
s: 46
i: [1,47,0,9,3,0,0,1] t59_orders_1
s: 47
i: [1,48,0,9,3,0,0,1] t59_orders_1
s: 48
i: [1,49,0,9,3,0,0,1] t51_orders_1
s: 49
i: [1,50,0,9,3,0,0,1] t51_orders_1
s: 50
d: [2,82,1,12,4,0,44,46,43]
d: [2,83,1,9,3,0,21,82,42]
s: 83
i: [1,45,0,10,2,0,0,255] d6_relat_1
s: 45
d: [2,84,1,12,4,0,44,47,43]
d: [2,85,1,9,3,0,21,84,42]
s: 85
i: [1,55,0,14,4,0,0,2] t63_orders_1
s: 55
d: [2,101,1,12,4,1,44,55,64]
d: [2,103,1,9,3,2,21,101,69]
d: [2,104,1,7,2,3,21,103,74]
a: [3,112,1,3,1,104]
s: 112
d: [2,87,1,12,4,0,44,48,43]
d: [2,89,1,9,3,0,21,87,42]
s: 89
d: [2,90,1,12,4,0,44,49,43]
d: [2,91,1,9,3,0,21,90,42]
s: 91
d: [2,93,1,12,4,0,44,50,43]
d: [2,95,1,9,3,0,21,93,42]
s: 95
d: [2,99,1,11,4,0,44,54,45]
s: 99
d: [2,102,1,21,5,0,44,55,44]
d: [2,114,1,18,4,0,21,102,42]
s: 114
d: [2,125,2,19,5,0,44,114,43]
d: [2,127,2,16,4,0,21,125,42]
s: 127
d: [2,124,2,19,5,0,40,114,48]
d: [2,126,2,17,4,0,18,124]
s: 126
d: [2,135,3,15,4,1,44,126,64]
d: [2,139,3,13,3,2,21,135,74]
d: [2,140,3,8,2,2,21,139,41]
a: [3,144,3,4,1,140]
s: 144
d: [2,151,4,0,0,3,52,74,69,144,46]
e: 151
d: [2,152,4,5,2,1,40,144,46]
d: [2,154,4,2,1,2,21,152,69]
d: [2,155,4,0,0,3,21,154,74]
e: 155
r: 144
d: [2,120,2,22,5,0,40,95,55]
d: [2,123,2,19,4,0,21,120,42]
s: 123
a: [3,143,3,4,1,140]
s: 143
a: [3,148,3,4,1,140]
s: 148
d: [2,164,4,0,0,3,52,74,59,148,95]
e: 164
d: [2,165,4,5,2,1,40,148,95]
d: [2,167,4,2,1,2,21,165,59]
d: [2,168,4,0,0,3,21,167,74]
e: 168
f: 168,155,140,37,38,39,40
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t64_orders_1
% SZS output start Proof for orders_1__t64_orders_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : ~(! [X2] : ~r7_orders_1(X0,X2) & r5_orders_1(X0,X1) & k1_relat_1(X0) = X1 & r2_orders_1(X0,X1))) [input t64_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : ~(! [X2] : ~r7_orders_1(X0,X2) & r5_orders_1(X0,X1) & k1_relat_1(X0) = X1 & r2_orders_1(X0,X1))) [negated conjecture 1]
7. ! [X0] : (v1_relat_1(X0) => v1_relat_1(k2_relat_1(X0))) [input dt_k2_relat_1]
12. ! [X0] : (v1_relat_1(X0) => k2_relat_1(k2_relat_1(X0)) = X0) [input involutiveness_k2_relat_1]
13. ! [X0] : (v1_relat_1(X0) => k1_relat_1(X0) = k1_relat_1(k2_relat_1(X0))) [input t21_relat_1]
14. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r2_orders_1(X0,X1) => r2_orders_1(k2_relat_1(X0),X1))) [input t41_orders_1]
15. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r4_orders_1(X0,X1) <=> r5_orders_1(k2_relat_1(X0),X1))) [input t51_orders_1]
16. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r7_orders_1(X0,X1) <=> r6_orders_1(k2_relat_1(X0),X1))) [input t59_orders_1]
17. ! [X0] : (v1_relat_1(X0) => ! [X1] : ~(! [X2] : ~r6_orders_1(X0,X2) & r4_orders_1(X0,X1) & k1_relat_1(X0) = X1 & r2_orders_1(X0,X1))) [input t63_orders_1]
19. ? [X0] : (? [X1] : (! [X2] : ~r7_orders_1(X0,X2) & r5_orders_1(X0,X1) & k1_relat_1(X0) = X1 & r2_orders_1(X0,X1)) & v1_relat_1(X0)) [ennf transformation 2]
20. ! [X0] : (v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 7]
21. ! [X0] : (k2_relat_1(k2_relat_1(X0)) = X0 | ~v1_relat_1(X0)) [ennf transformation 12]
22. ! [X0] : (k1_relat_1(X0) = k1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 13]
24. ! [X0] : (! [X1] : (r2_orders_1(k2_relat_1(X0),X1) | ~r2_orders_1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 14]
25. ! [X0] : (! [X1] : (r7_orders_1(X0,X1) <=> r6_orders_1(k2_relat_1(X0),X1)) | ~v1_relat_1(X0)) [ennf transformation 16]
26. ! [X0] : (! [X1] : (r4_orders_1(X0,X1) <=> r5_orders_1(k2_relat_1(X0),X1)) | ~v1_relat_1(X0)) [ennf transformation 15]
27. ! [X0] : (! [X1] : (? [X2] : r6_orders_1(X0,X2) | ~r4_orders_1(X0,X1) | k1_relat_1(X0) != X1 | ~r2_orders_1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 17]
30. ? [X0] : (? [X1] : (! [X2] : ~r7_orders_1(X0,X2) & r5_orders_1(X0,X1) & k1_relat_1(X0) = X1 & r2_orders_1(X0,X1)) & v1_relat_1(X0)) => (? [X1] : (! [X2] : ~r7_orders_1(sK0,X2) & r5_orders_1(sK0,X1) & k1_relat_1(sK0) = X1 & r2_orders_1(sK0,X1)) & v1_relat_1(sK0)) [choice axiom]
31. ? [X1] : (! [X2] : ~r7_orders_1(sK0,X2) & r5_orders_1(sK0,X1) & k1_relat_1(sK0) = X1 & r2_orders_1(sK0,X1)) => (! [X2] : ~r7_orders_1(sK0,X2) & r5_orders_1(sK0,sK1) & k1_relat_1(sK0) = sK1 & r2_orders_1(sK0,sK1)) [choice axiom]
32. (! [X2] : ~r7_orders_1(sK0,X2) & r5_orders_1(sK0,sK1) & k1_relat_1(sK0) = sK1 & r2_orders_1(sK0,sK1)) & v1_relat_1(sK0) [skolemisation 19,31,30]
33. ! [X0] : (! [X1] : ((r7_orders_1(X0,X1) | ~r6_orders_1(k2_relat_1(X0),X1)) & (r6_orders_1(k2_relat_1(X0),X1) | ~r7_orders_1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 25]
34. ! [X0] : (! [X1] : ((r4_orders_1(X0,X1) | ~r5_orders_1(k2_relat_1(X0),X1)) & (r5_orders_1(k2_relat_1(X0),X1) | ~r4_orders_1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 26]
35. ! [X0] : (? [X2] : r6_orders_1(X0,X2) => r6_orders_1(X0,sK2(X0))) [choice axiom]
36. ! [X0] : (! [X1] : (r6_orders_1(X0,sK2(X0)) | ~r4_orders_1(X0,X1) | k1_relat_1(X0) != X1 | ~r2_orders_1(X0,X1)) | ~v1_relat_1(X0)) [skolemisation 27,35]
37. v1_relat_1(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. r2_orders_1(sK0,sK1) [cnf transformation 32] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. k1_relat_1(sK0) = sK1 [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. r5_orders_1(sK0,sK1) [cnf transformation 32] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. ~r7_orders_1(sK0,X2) [cnf transformation 32] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 20] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
43. k2_relat_1(k2_relat_1(X0)) = X0 | ~v1_relat_1(X0) [cnf transformation 21] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
44. k1_relat_1(X0) = k1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
46. r2_orders_1(k2_relat_1(X0),X1) | ~r2_orders_1(X0,X1) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
48. ~r6_orders_1(k2_relat_1(X0),X1) | r7_orders_1(X0,X1) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. ~r5_orders_1(k2_relat_1(X0),X1) | r4_orders_1(X0,X1) | ~v1_relat_1(X0) [cnf transformation 34] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. r6_orders_1(X0,sK2(X0)) | ~r4_orders_1(X0,X1) | k1_relat_1(X0) != X1 | ~r2_orders_1(X0,X1) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:16,wCS:6400,sine:2,thAx:0,allAx:2,thDist:-2}
55. ~r4_orders_1(X0,k1_relat_1(X0)) | r6_orders_1(X0,sK2(X0)) | ~r2_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) [equality resolution 51] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
57. 1 <=> r5_orders_1(sK0,sK1) [avatar definition]
59. r5_orders_1(sK0,sK1) <- (1) [avatar component clause 57] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
60. 1 [avatar split clause 40,57]
62. 2 <=> k1_relat_1(sK0) = sK1 [avatar definition]
64. k1_relat_1(sK0) = sK1 <- (2) [avatar component clause 62] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
65. 2 [avatar split clause 39,62]
67. 3 <=> r2_orders_1(sK0,sK1) [avatar definition]
69. r2_orders_1(sK0,sK1) <- (3) [avatar component clause 67] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
70. 3 [avatar split clause 38,67]
72. 4 <=> v1_relat_1(sK0) [avatar definition]
74. v1_relat_1(sK0) <- (4) [avatar component clause 72] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
75. 4 [avatar split clause 37,72]
93. ~r5_orders_1(X0,X1) | r4_orders_1(k2_relat_1(X0),X1) | ~v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [superposition 50,43] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
95. r4_orders_1(k2_relat_1(X0),X1) | ~r5_orders_1(X0,X1) | ~v1_relat_1(X0) [subsumption resolution 93,42] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
102. ~r4_orders_1(k2_relat_1(X0),k1_relat_1(X0)) | r6_orders_1(k2_relat_1(X0),sK2(k2_relat_1(X0))) | ~r2_orders_1(k2_relat_1(X0),k1_relat_1(X0)) | ~v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [superposition 55,44] {a:1,w:21,wCS:8400,sine:1,thAx:0,allAx:3,thDist:-3}
114. r6_orders_1(k2_relat_1(X0),sK2(k2_relat_1(X0))) | ~r4_orders_1(k2_relat_1(X0),k1_relat_1(X0)) | ~r2_orders_1(k2_relat_1(X0),k1_relat_1(X0)) | ~v1_relat_1(X0) [subsumption resolution 102,42] {a:1,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
124. ~r4_orders_1(k2_relat_1(X0),k1_relat_1(X0)) | ~r2_orders_1(k2_relat_1(X0),k1_relat_1(X0)) | ~v1_relat_1(X0) | r7_orders_1(X0,sK2(k2_relat_1(X0))) | ~v1_relat_1(X0) [resolution 114,48] {a:2,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
126. ~r4_orders_1(k2_relat_1(X0),k1_relat_1(X0)) | ~r2_orders_1(k2_relat_1(X0),k1_relat_1(X0)) | ~v1_relat_1(X0) | r7_orders_1(X0,sK2(k2_relat_1(X0))) [duplicate literal removal 124] {a:2,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
135. ~r4_orders_1(k2_relat_1(sK0),sK1) | ~r2_orders_1(k2_relat_1(sK0),sK1) | ~v1_relat_1(sK0) | r7_orders_1(sK0,sK2(k2_relat_1(sK0))) <- (2) [superposition 126,64] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
139. ~r4_orders_1(k2_relat_1(sK0),sK1) | ~r2_orders_1(k2_relat_1(sK0),sK1) | r7_orders_1(sK0,sK2(k2_relat_1(sK0))) <- (2, 4) [subsumption resolution 135,74] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
140. ~r4_orders_1(k2_relat_1(sK0),sK1) | ~r2_orders_1(k2_relat_1(sK0),sK1) <- (2, 4) [subsumption resolution 139,41] {a:3,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
142. 7 <=> r2_orders_1(k2_relat_1(sK0),sK1) [avatar definition]
144. ~r2_orders_1(k2_relat_1(sK0),sK1) <- (~7) [avatar component clause 142] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
146. 8 <=> r4_orders_1(k2_relat_1(sK0),sK1) [avatar definition]
148. ~r4_orders_1(k2_relat_1(sK0),sK1) <- (~8) [avatar component clause 146] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
149. ~7 | ~8 | ~2 | ~4 [avatar split clause 140,72,62,146,142]
152. ~r2_orders_1(sK0,sK1) | ~v1_relat_1(sK0) <- (~7) [resolution 144,46] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
154. ~v1_relat_1(sK0) <- (3, ~7) [subsumption resolution 152,69] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
155. $false <- (3, 4, ~7) [subsumption resolution 154,74] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
156. ~3 | ~4 | 7 [avatar contradiction clause 155]
165. ~r5_orders_1(sK0,sK1) | ~v1_relat_1(sK0) <- (~8) [resolution 148,95] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
167. ~v1_relat_1(sK0) <- (1, ~8) [subsumption resolution 165,59] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
168. $false <- (1, 4, ~8) [subsumption resolution 167,74] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
169. ~1 | ~4 | 8 [avatar contradiction clause 168]
170. $false [avatar sat refutation 60,65,70,75,149,156,169] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:83,thDist:-83}
% SZS output end Proof for orders_1__t64_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 18
% Generated clauses: 98
% Active clauses: 32
% Passive clauses: 39
% Final active clauses: 31
% Final passive clauses: 7
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 20
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 2
% Forward superposition: 24
% Backward superposition: 2
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 4
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward superposition index maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

