i: [1,90,0,8,2,1,0,0] t20_mcart_1
a: [3,145,0,4,1,90]
s: 145
i: [1,89,0,5,1,1,0,0] t20_mcart_1
a: [3,150,0,5,1,89]
s: 150
i: [1,91,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,155,0,2,1,91]
s: 155
i: [1,92,0,2,1,0,0,255] fc1_xboole_0
a: [3,160,0,2,1,92]
s: 160
i: [1,124,0,2,1,0,0,2] rc2_xboole_0
a: [3,170,0,2,1,124]
s: 170
i: [1,125,0,2,1,0,0,2] rc1_xtuple_0
a: [3,175,0,2,1,125]
s: 175
i: [1,126,0,2,1,0,0,2] rc1_xboole_0
a: [3,180,0,2,1,126]
s: 180
i: [1,128,0,2,1,0,0,255] rc1_relat_1
a: [3,185,0,2,1,128]
s: 185
i: [1,127,0,2,1,0,0,255] rc1_relat_1
a: [3,190,0,2,1,127]
s: 190
i: [1,129,0,2,1,0,0,255] rc2_relat_1
a: [3,195,0,2,1,129]
s: 195
i: [1,94,0,3,1,0,0,255] fc2_xboole_0
s: 94
i: [1,93,0,3,1,0,0,255] d2_xboole_0
a: [3,165,0,3,1,93]
s: 165
r: 155
i: [1,96,0,4,2,0,0,255] cc1_relat_1
s: 96
d: [2,198,1,2,1,1,40,96,180]
a: [3,203,1,2,1,198]
s: 203
i: [1,115,0,4,1,0,0,255] existence_m1_subset_1
s: 115
i: [1,116,0,4,1,0,0,255] fc3_xboole_0
s: 116
d: [2,199,1,2,1,1,40,96,160]
a: [3,208,1,2,1,199]
s: 208
i: [1,117,0,4,1,0,0,2] fc1_xtuple_0
s: 117
d: [2,210,1,2,1,1,44,117,150]
a: [3,214,1,2,1,210]
s: 214
i: [1,131,0,4,1,0,0,255] d1_tarski
s: 131
i: [1,97,0,5,2,0,0,255] t6_boole
s: 97
i: [1,108,0,5,2,0,0,255] t7_boole
s: 108
i: [1,118,0,5,1,0,0,255] fc5_relat_1
s: 118
d: [2,220,1,3,1,1,44,118,150]
a: [3,224,1,3,1,220]
s: 224
i: [1,134,0,5,1,0,0,255] d2_tarski
s: 134
i: [1,136,0,5,1,0,0,255] d2_tarski
s: 136
i: [1,101,0,6,2,0,0,255] t1_subset
s: 101
i: [1,119,0,6,1,0,0,1] rd1_xtuple_0
s: 119
d: [2,231,1,4,1,1,44,119,150]
a: [3,235,1,4,1,231]
s: 235
i: [1,120,0,6,1,0,0,1] rd2_xtuple_0
s: 120
i: [1,122,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 122
d: [2,237,1,4,1,1,44,120,150]
d: [2,238,1,3,1,2,22,237,145]
a: [3,242,1,3,1,238]
s: 242
r: 150
d: [2,247,0,5,1,2,23,150,242]
a: [3,251,0,5,1,247]
s: 251
d: [2,228,1,4,1,0,40,101,131]
s: 228
i: [1,107,0,7,3,0,0,1] t8_boole
s: 107
i: [1,121,0,7,1,0,0,255] commutativity_k2_tarski
s: 121
d: [2,244,1,4,1,0,40,122,131]
s: 244
i: [1,132,0,7,2,0,0,255] d1_tarski
s: 132
d: [2,229,1,5,1,0,40,101,134]
s: 229
i: [1,102,0,8,3,0,0,255] t2_subset
s: 102
i: [1,123,0,8,1,0,0,255] fc7_relat_1
s: 123
d: [2,269,1,6,1,1,44,123,251]
s: 269
i: [1,95,0,9,2,0,0,1] rd3_xtuple_0
s: 95
d: [2,274,2,4,1,1,44,269,251]
a: [3,278,2,4,1,274]
s: 278
i: [1,100,0,10,1,0,0,255] d5_tarski
s: 100
i: [1,137,0,11,3,0,0,255] d2_tarski
s: 137
d: [2,270,1,6,1,1,44,123,251]
d: [2,273,1,6,1,1,22,270,121]
s: 273
i: [1,105,0,14,3,0,0,255] d1_tarski
s: 105
d: [2,288,2,6,2,1,44,269,95]
s: 288
i: [1,106,0,14,3,0,0,255] d1_tarski
s: 106
i: [1,113,0,17,3,0,0,255] d2_tarski
s: 113
d: [2,316,3,6,2,1,44,288,121]
s: 316
i: [1,114,0,17,3,0,0,255] d2_tarski
s: 114
d: [2,317,3,6,2,1,44,288,121]
s: 317
i: [1,112,0,23,4,0,0,255] d2_tarski
s: 112
d: [2,300,1,4,1,0,44,116,100]
s: 300
d: [2,325,2,2,1,1,44,300,251]
a: [3,330,2,2,1,325]
s: 330
d: [2,230,1,5,1,0,40,101,136]
s: 230
d: [2,326,2,4,2,0,44,300,95]
s: 326
d: [2,245,1,5,1,0,40,122,134]
s: 245
d: [2,246,1,5,1,0,40,122,136]
s: 246
d: [2,336,3,2,1,1,40,326,175]
a: [3,342,3,2,1,336]
s: 342
d: [2,285,1,5,2,0,44,118,95]
s: 285
d: [2,264,2,5,1,0,44,229,121]
s: 264
d: [2,266,1,6,2,0,40,102,115]
s: 266
d: [2,303,1,6,1,0,44,134,100]
s: 303
r: 145
r: 242
r: 316
r: 288
r: 269
r: 251
a: [3,241,1,3,1,238]
s: 241
a: [3,141,0,4,1,90]
s: 141
r: 235
r: 141
r: 241
s: 242
a: [3,140,0,4,1,90]
s: 140
s: 145
s: 235
s: 150
d: [2,377,1,3,1,2,23,235,141]
a: [3,381,1,3,1,377]
s: 381
s: 251
a: [3,375,1,4,1,237]
s: 375
d: [2,360,2,4,1,1,44,303,251]
a: [3,365,2,4,1,360]
s: 365
r: 140
r: 145
r: 242
r: 251
r: 381
s: 141
r: 235
s: 241
a: [3,144,0,4,1,90]
s: 144
s: 150
a: [3,382,1,3,1,377]
s: 382
r: 150
r: 365
d: [2,449,2,4,1,2,23,365,382]
d: [2,455,2,0,0,2,21,449,244]
e: 455
d: [2,443,2,4,1,1,44,303,150]
d: [2,457,2,4,1,2,22,443,382]
d: [2,458,2,0,0,2,21,457,244]
e: 458
r: 141
r: 144
r: 241
r: 382
s: 140
s: 242
s: 145
s: 150
s: 381
s: 251
s: 235
d: [2,448,0,5,1,2,23,150,382]
a: [3,453,0,5,1,448]
s: 453
d: [2,483,0,5,1,2,22,453,242]
a: [3,487,0,5,1,483]
s: 487
d: [2,491,1,4,1,1,44,119,487]
d: [2,497,1,0,0,2,21,491,140]
e: 497
d: [2,496,2,4,1,1,44,303,487]
d: [2,500,2,0,0,1,21,496,244]
e: 500
r: 487
s: 375
a: [3,486,0,5,1,483]
s: 486
s: 365
d: [2,390,1,8,2,1,44,95,145]
d: [2,391,1,6,1,2,21,390,214]
a: [3,395,1,6,1,391]
s: 395
d: [2,440,1,6,1,1,44,123,150]
s: 440
d: [2,416,3,4,1,1,40,365,122]
a: [3,422,3,4,1,416]
s: 422
d: [2,441,1,6,1,1,44,123,150]
d: [2,447,1,6,1,1,22,441,121]
s: 447
d: [2,417,3,4,1,1,40,365,101]
a: [3,427,3,4,1,417]
s: 427
d: [2,472,1,6,1,1,44,123,251]
s: 472
d: [2,473,1,6,1,1,44,123,251]
d: [2,477,1,6,1,1,22,473,121]
s: 477
d: [2,281,2,8,2,1,44,95,235]
d: [2,290,2,6,1,2,21,281,214]
a: [3,370,2,6,1,290]
s: 370
d: [2,216,1,7,3,0,44,97,97]
s: 216
d: [2,514,2,6,2,1,44,440,95]
s: 514
d: [2,304,1,7,1,0,44,136,100]
s: 304
d: [2,527,2,5,1,1,44,304,251]
d: [2,528,2,0,0,1,21,527,245]
e: 528
d: [2,438,1,4,1,1,44,119,150]
d: [2,444,1,3,1,2,22,438,141]
d: [2,466,0,5,1,2,22,150,242]
f: 528,466,458,444,238,89,90
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t20_mcart_1
% SZS output start Proof for mcart_1__t20_mcart_1
1. ! [X0] : (? [X1] : ? [X2] : k4_tarski(X1,X2) = X0 => (k2_xtuple_0(X0) != X0 & k1_xtuple_0(X0) != X0)) [input t20_mcart_1]
2. ~! [X0] : (? [X1] : ? [X2] : k4_tarski(X1,X2) = X0 => (k2_xtuple_0(X0) != X0 & k1_xtuple_0(X0) != X0)) [negated conjecture 1]
3. ! [X0,X1] : (r2_hidden(X0,X1) => ~r2_hidden(X1,X0)) [input antisymmetry_r2_hidden]
8. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
9. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
11. ! [X0] : ! [X1] : k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [input d5_tarski]
32. ! [X0,X1] : k1_xtuple_0(k4_tarski(X0,X1)) = X0 [input rd1_xtuple_0]
33. ! [X0,X1] : k2_xtuple_0(k4_tarski(X0,X1)) = X1 [input rd2_xtuple_0]
40. ~! [X0] : (? [X1,X2] : k4_tarski(X1,X2) = X0 => (k2_xtuple_0(X0) != X0 & k1_xtuple_0(X0) != X0)) [flattening 2]
41. ! [X0,X1] : k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [flattening 11]
44. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 8]
47. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 9]
51. ? [X0] : ((k2_xtuple_0(X0) = X0 | k1_xtuple_0(X0) = X0) & ? [X1,X2] : k4_tarski(X1,X2) = X0) [ennf transformation 40]
64. ! [X0,X1] : (~r2_hidden(X1,X0) | ~r2_hidden(X0,X1)) [ennf transformation 3]
65. ? [X0] : ((k2_xtuple_0(X0) = X0 | k1_xtuple_0(X0) = X0) & ? [X1,X2] : k4_tarski(X1,X2) = X0) => ((sK0 = k2_xtuple_0(sK0) | sK0 = k1_xtuple_0(sK0)) & ? [X2,X1] : k4_tarski(X1,X2) = sK0) [choice axiom]
66. ? [X2,X1] : k4_tarski(X1,X2) = sK0 => sK0 = k4_tarski(sK1,sK2) [choice axiom]
67. (sK0 = k2_xtuple_0(sK0) | sK0 = k1_xtuple_0(sK0)) & sK0 = k4_tarski(sK1,sK2) [skolemisation 51,66,65]
68. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 44]
69. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 68]
70. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
71. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 69,70]
72. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 47]
73. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 72]
74. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 73]
75. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
76. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 74,75]
89. sK0 = k4_tarski(sK1,sK2) [cnf transformation 67] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
90. sK0 = k2_xtuple_0(sK0) | sK0 = k1_xtuple_0(sK0) [cnf transformation 67] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [cnf transformation 41] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
104. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 71] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
110. r2_hidden(X4,X2) | X0 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 76] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
111. r2_hidden(X4,X2) | X1 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 76] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
119. k1_xtuple_0(k4_tarski(X0,X1)) = X0 [cnf transformation 32] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. k2_xtuple_0(k4_tarski(X0,X1)) = X1 [cnf transformation 33] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
122. ~r2_hidden(X1,X0) | ~r2_hidden(X0,X1) [cnf transformation 64] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
130. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 104] {a:0,w:7,wCS:2800,sine:255,thAx:0,allAx:2,thDist:-2}
131. r2_hidden(X3,k1_tarski(X3)) [equality resolution 130] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
133. r2_hidden(X4,X2) | k2_tarski(X0,X4) != X2 [equality resolution 111] {a:0,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
134. r2_hidden(X4,k2_tarski(X0,X4)) [equality resolution 133] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
135. r2_hidden(X4,X2) | k2_tarski(X4,X1) != X2 [equality resolution 110] {a:0,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
136. r2_hidden(X4,k2_tarski(X4,X1)) [equality resolution 135] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
139. 1 <=> sK0 = k1_xtuple_0(sK0) [avatar definition]
141. sK0 = k1_xtuple_0(sK0) <- (1) [avatar component clause 139] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
143. 2 <=> sK0 = k2_xtuple_0(sK0) [avatar definition]
145. sK0 = k2_xtuple_0(sK0) <- (2) [avatar component clause 143] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
146. 1 | 2 [avatar split clause 90,143,139]
148. 3 <=> sK0 = k4_tarski(sK1,sK2) [avatar definition]
150. sK0 = k4_tarski(sK1,sK2) <- (3) [avatar component clause 148] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
151. 3 [avatar split clause 89,148]
237. k2_xtuple_0(sK0) = sK2 <- (3) [superposition 120,150] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
238. sK0 = sK2 <- (2, 3) [forward demodulation 237,145] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
240. 18 <=> sK0 = sK2 [avatar definition]
242. sK0 = sK2 <- (18) [avatar component clause 240] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
243. 18 | ~2 | ~3 [avatar split clause 238,148,143,240]
244. ~r2_hidden(k1_tarski(X0),X0) [resolution 122,131] {a:1,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
245. ~r2_hidden(k2_tarski(X1,X2),X2) [resolution 122,134] {a:1,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
249. 19 <=> sK0 = k4_tarski(sK1,sK0) [avatar definition]
251. sK0 = k4_tarski(sK1,sK0) <- (19) [avatar component clause 249] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
303. r2_hidden(k1_tarski(X6),k4_tarski(X6,X7)) [superposition 134,100] {a:1,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
304. r2_hidden(k2_tarski(X8,X9),k4_tarski(X8,X9)) [superposition 136,100] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
380. 26 <=> sK0 = sK1 [avatar definition]
382. sK0 = sK1 <- (26) [avatar component clause 380] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
438. k1_xtuple_0(sK0) = sK1 <- (3) [superposition 119,150] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
443. r2_hidden(k1_tarski(sK1),sK0) <- (3) [superposition 303,150] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
444. sK0 = sK1 <- (1, 3) [forward demodulation 438,141] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
445. 26 | ~1 | ~3 [avatar split clause 444,148,139,380]
457. r2_hidden(k1_tarski(sK0),sK0) <- (3, 26) [forward demodulation 443,382] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
458. $false <- (3, 26) [subsumption resolution 457,244] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
459. ~3 | ~26 [avatar contradiction clause 458]
466. sK0 = k4_tarski(sK1,sK0) <- (3, 18) [forward demodulation 150,242] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
467. 19 | ~3 | ~18 [avatar split clause 466,240,148,249]
527. r2_hidden(k2_tarski(sK1,sK0),sK0) <- (19) [superposition 304,251] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
528. $false <- (19) [subsumption resolution 527,245] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
529. ~19 [avatar contradiction clause 528]
530. $false [avatar sat refutation 146,151,243,445,459,467,529] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:339,thDist:-339}
% SZS output end Proof for mcart_1__t20_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 314
% Active clauses: 91
% Passive clauses: 152
% Final active clauses: 67
% Final passive clauses: 41
% Main loop iterations started: 105
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 17
% Fw demodulations: 22
% Bw demodulations: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Equational tautologies: 3
% Forward subsumptions: 69
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Binary resolution: 49
% Forward superposition: 48
% Backward superposition: 63
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 59
% SAT solver unit clauses: 12
% SAT solver binary clauses: 24
% 
% Memory used [KB]: 5117
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% congruence closure: 0.001 s
% superposition: 0.001 s
% other: 0.005 s ( own 0.001 s ) 

