i: [1,70,0,4,2,1,0,0] t13_yellow_0
a: [3,96,0,2,1,70]
s: 96
i: [1,69,0,4,2,1,0,0] t13_yellow_0
a: [3,105,0,2,1,69]
s: 105
i: [1,65,0,2,1,1,0,0] t13_yellow_0
a: [3,121,0,2,1,65]
s: 121
i: [1,64,0,2,1,1,0,0] t13_yellow_0
a: [3,126,0,2,1,64]
s: 126
i: [1,88,0,2,1,0,0,1] existence_l1_orders_2
a: [3,131,0,2,1,88]
s: 131
i: [1,90,0,2,1,0,0,255] rc1_orders_2
a: [3,136,0,2,1,90]
s: 136
i: [1,89,0,2,1,0,0,255] rc1_orders_2
a: [3,141,0,2,1,89]
s: 141
i: [1,66,0,11,1,1,0,0] t13_yellow_0
a: [3,116,0,11,1,66]
s: 116
i: [1,83,0,3,1,0,0,255] reflexivity_r1_tarski
s: 83
i: [1,82,0,4,1,0,0,2] existence_m1_subset_1
s: 82
i: [1,81,0,7,2,0,0,255] t3_subset
s: 81
i: [1,73,0,9,3,0,0,1] d4_yellow_0
s: 73
i: [1,76,0,9,3,0,0,1] d5_yellow_0
s: 76
i: [1,74,0,10,3,0,0,1] d4_yellow_0
s: 74
i: [1,77,0,10,3,0,0,1] d5_yellow_0
s: 77
i: [1,84,0,10,2,0,0,255] dt_g1_orders_2
s: 84
d: [2,144,1,8,1,0,40,84,82]
s: 144
i: [1,85,0,10,2,0,0,255] dt_g1_orders_2
s: 85
d: [2,146,1,8,1,0,40,85,82]
s: 146
i: [1,71,0,11,2,0,0,1] dt_u1_orders_2
s: 71
i: [1,72,0,11,3,0,0,255] abstractness_v1_orders_2
s: 72
d: [2,149,1,11,3,1,44,72,116]
d: [2,151,1,9,2,2,21,149,121]
a: [3,155,1,2,1,151]
s: 155
i: [1,75,0,13,4,0,0,1] d4_yellow_0
s: 75
d: [2,147,1,8,2,0,40,71,85]
s: 147
i: [1,78,0,13,4,0,0,1] d5_yellow_0
s: 78
i: [1,86,0,16,3,0,0,1] free_g1_orders_2
s: 86
d: [2,165,2,8,2,1,44,147,116]
d: [2,168,2,6,1,2,21,165,121]
a: [3,172,2,6,1,168]
s: 172
i: [1,87,0,16,3,0,0,1] free_g1_orders_2
s: 87
d: [2,176,1,22,3,1,44,86,116]
a: [3,184,1,9,1,176]
s: 184
d: [2,201,2,0,0,2,52,121,184,71]
e: 201
d: [2,202,2,2,1,1,40,184,71]
d: [2,205,2,0,0,2,21,202,121]
e: 205
r: 184
i: [1,91,0,31,7,0,0,2] t2_yellow_0
s: 91
i: [1,92,0,31,7,0,0,2] t2_yellow_0
s: 92
a: [3,183,1,9,1,176]
s: 183
a: [3,187,1,13,2,176]
s: 187
d: [2,238,2,13,4,1,44,187,72]
d: [2,240,2,11,3,2,21,238,126]
a: [3,244,2,2,1,240]
s: 244
d: [2,191,1,22,3,1,44,87,116]
a: [3,198,1,13,2,191]
s: 198
d: [2,178,1,19,3,1,44,86,116]
s: 178
d: [2,239,2,5,1,1,47,187]
a: [3,252,2,5,1,239]
s: 252
r: 187
r: 183
r: 116
a: [3,99,0,2,1,70]
s: 99
d: [2,257,2,5,1,1,47,198]
a: [3,261,2,5,1,257]
s: 261
r: 198
d: [2,265,0,11,1,2,23,116,252]
a: [3,287,0,11,1,265]
s: 287
d: [2,264,1,9,1,2,23,183,252]
a: [3,282,1,9,1,264]
s: 282
d: [2,271,3,9,3,1,44,76,252]
d: [2,293,3,7,2,2,21,271,121]
a: [3,297,3,5,1,293]
s: 297
d: [2,263,1,13,2,2,23,187,252]
s: 263
d: [2,228,2,6,1,1,40,183,84]
d: [2,230,2,6,1,2,22,228,116]
a: [3,234,2,6,1,230]
s: 234
d: [2,320,1,13,2,2,23,198,261]
s: 320
d: [2,193,1,19,3,1,44,87,116]
s: 193
d: [2,272,3,10,3,1,44,77,252]
d: [2,299,3,8,2,2,21,272,121]
a: [3,303,3,6,1,299]
s: 303
d: [2,207,1,31,7,1,44,91,116]
d: [2,215,1,29,6,2,21,207,121]
s: 215
d: [2,278,3,8,2,1,44,147,252]
d: [2,314,3,6,1,2,21,278,121]
a: [3,318,3,6,1,314]
s: 318
d: [2,209,1,31,7,1,44,91,116]
d: [2,216,1,29,6,2,21,209,121]
s: 216
d: [2,217,1,31,7,1,44,92,116]
d: [2,225,1,29,6,2,21,217,121]
s: 225
d: [2,270,3,13,4,1,44,75,252]
d: [2,291,3,11,3,2,21,270,96]
d: [2,292,3,9,2,3,21,291,121]
s: 292
d: [2,219,1,31,7,1,44,92,116]
d: [2,226,1,29,6,2,21,219,121]
s: 226
d: [2,321,3,11,2,1,44,71,261]
d: [2,328,3,9,1,2,21,321,121]
d: [2,329,3,9,1,3,22,328,252]
a: [3,333,3,9,1,329]
s: 333
d: [2,359,1,29,6,3,22,215,252]
s: 359
d: [2,361,1,29,6,3,22,216,252]
s: 361
d: [2,370,2,18,5,3,47,359]
d: [2,371,2,14,4,3,18,370]
d: [2,374,2,12,3,4,21,371,126]
s: 374
d: [2,362,1,29,6,3,22,225,252]
s: 362
d: [2,379,2,18,5,3,47,361]
d: [2,380,2,14,4,3,18,379]
d: [2,383,2,12,3,4,21,380,126]
s: 383
d: [2,363,1,29,6,3,22,226,252]
s: 363
d: [2,148,1,8,2,0,40,71,84]
s: 148
d: [2,394,3,15,4,4,40,383,74]
d: [2,395,3,13,3,5,21,394,105]
d: [2,396,3,11,2,5,21,395,126]
d: [2,397,3,5,1,6,21,396,292]
a: [3,401,3,5,1,397]
s: 401
d: [2,420,4,0,0,3,52,126,105,401,73]
e: 420
d: [2,421,4,4,2,1,40,401,73]
d: [2,423,4,2,1,2,21,421,105]
d: [2,424,4,0,0,3,21,423,126]
e: 424
r: 99
r: 105
a: [3,100,0,2,1,70]
s: 100
a: [3,104,0,2,1,69]
s: 104
i: [1,68,0,4,2,1,0,0] t13_yellow_0
a: [3,110,0,2,1,68]
s: 110
d: [2,143,1,9,2,0,40,84,81]
s: 143
d: [2,273,3,13,4,1,44,78,252]
d: [2,305,3,11,3,2,21,273,121]
a: [3,308,3,9,2,305]
s: 308
d: [2,429,4,0,0,3,52,297,303,308]
e: 429
d: [2,430,4,5,1,2,40,308,303]
d: [2,432,4,0,0,3,21,430,297]
e: 432
r: 297
d: [2,145,1,9,2,0,40,85,81]
s: 145
a: [3,296,3,5,1,293]
s: 296
d: [2,177,1,22,5,0,44,86,72]
d: [2,189,1,13,4,0,21,177,71]
s: 189
d: [2,192,1,22,5,0,44,87,72]
d: [2,200,1,13,4,0,21,192,71]
s: 200
d: [2,389,2,18,5,3,47,362]
d: [2,390,2,14,4,3,18,389]
d: [2,393,2,12,3,4,21,390,126]
s: 393
d: [2,179,1,19,5,0,44,86,72]
s: 179
d: [2,407,2,18,5,3,47,363]
d: [2,408,2,14,4,3,18,407]
d: [2,411,2,12,3,4,21,408,126]
s: 411
d: [2,194,1,19,5,0,44,87,72]
s: 194
d: [2,210,1,31,9,0,44,91,72]
d: [2,213,1,29,8,0,18,210]
s: 213
d: [2,442,3,15,4,4,40,411,77]
d: [2,443,3,13,3,5,21,442,110]
d: [2,444,3,11,2,5,21,443,126]
d: [2,445,3,5,1,6,21,444,308]
a: [3,449,3,5,1,445]
s: 449
d: [2,463,4,0,0,3,52,126,110,449,76]
e: 463
d: [2,464,4,4,2,1,40,449,76]
d: [2,466,4,2,1,2,21,464,110]
d: [2,467,4,0,0,3,21,466,126]
e: 467
i: [1,67,0,4,2,1,0,0] t13_yellow_0
d: [2,426,4,2,1,2,21,421,126]
f: 467,445,426,397,305,239,205,176,64,65,66,67,68,69,70
% Refutation found. Thanks to Tanya!
% SZS status Theorem for yellow_0__t13_yellow_0
% SZS output start Proof for yellow_0__t13_yellow_0
1. ! [X0] : (l1_orders_2(X0) => ! [X1] : (l1_orders_2(X1) => (g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) => ((v2_yellow_0(X0) => v2_yellow_0(X1)) & (v1_yellow_0(X0) => v1_yellow_0(X1)))))) [input t13_yellow_0]
2. ~! [X0] : (l1_orders_2(X0) => ! [X1] : (l1_orders_2(X1) => (g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) => ((v2_yellow_0(X0) => v2_yellow_0(X1)) & (v1_yellow_0(X0) => v1_yellow_0(X1)))))) [negated conjecture 1]
6. ! [X0] : (l1_orders_2(X0) => (v1_yellow_0(X0) <=> ? [X1] : (r1_lattice3(X0,u1_struct_0(X0),X1) & m1_subset_1(X1,u1_struct_0(X0))))) [input d4_yellow_0]
7. ! [X0] : (l1_orders_2(X0) => (v2_yellow_0(X0) <=> ? [X1] : (r2_lattice3(X0,u1_struct_0(X0),X1) & m1_subset_1(X1,u1_struct_0(X0))))) [input d5_yellow_0]
14. ! [X0] : (l1_orders_2(X0) => m1_subset_1(u1_orders_2(X0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))))) [input dt_u1_orders_2]
19. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) => ! [X2,X3] : (g1_orders_2(X0,X1) = g1_orders_2(X2,X3) => (X1 = X3 & X0 = X2))) [input free_g1_orders_2]
22. ! [X0] : (l1_orders_2(X0) => ! [X1] : (l1_orders_2(X1) => (g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) => ! [X2] : ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X1)) => (X3 = X4 => ((r1_lattice3(X0,X2,X3) => r1_lattice3(X1,X2,X4)) & (r2_lattice3(X0,X2,X3) => r2_lattice3(X1,X2,X4))))))))) [input t2_yellow_0]
24. ! [X0] : (l1_orders_2(X0) => ! [X1] : (l1_orders_2(X1) => (g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) => ! [X2,X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X1)) => (X3 = X4 => ((r1_lattice3(X0,X2,X3) => r1_lattice3(X1,X2,X4)) & (r2_lattice3(X0,X2,X3) => r2_lattice3(X1,X2,X4))))))))) [flattening 22]
35. ? [X0] : (? [X1] : ((((~v2_yellow_0(X1) & v2_yellow_0(X0)) | (~v1_yellow_0(X1) & v1_yellow_0(X0))) & g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))) & l1_orders_2(X1)) & l1_orders_2(X0)) [ennf transformation 2]
36. ? [X0] : (? [X1] : (((~v2_yellow_0(X1) & v2_yellow_0(X0)) | (~v1_yellow_0(X1) & v1_yellow_0(X0))) & g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) & l1_orders_2(X1)) & l1_orders_2(X0)) [flattening 35]
37. ! [X0] : (m1_subset_1(u1_orders_2(X0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0)))) | ~l1_orders_2(X0)) [ennf transformation 14]
40. ! [X0] : ((v1_yellow_0(X0) <=> ? [X1] : (r1_lattice3(X0,u1_struct_0(X0),X1) & m1_subset_1(X1,u1_struct_0(X0)))) | ~l1_orders_2(X0)) [ennf transformation 6]
41. ! [X0] : ((v2_yellow_0(X0) <=> ? [X1] : (r2_lattice3(X0,u1_struct_0(X0),X1) & m1_subset_1(X1,u1_struct_0(X0)))) | ~l1_orders_2(X0)) [ennf transformation 7]
42. ! [X0] : (! [X1] : ((! [X2,X3] : (! [X4] : ((((r1_lattice3(X1,X2,X4) | ~r1_lattice3(X0,X2,X3)) & (r2_lattice3(X1,X2,X4) | ~r2_lattice3(X0,X2,X3))) | X3 != X4) | ~m1_subset_1(X4,u1_struct_0(X1))) | ~m1_subset_1(X3,u1_struct_0(X0))) | g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))) | ~l1_orders_2(X1)) | ~l1_orders_2(X0)) [ennf transformation 24]
43. ! [X0] : (! [X1] : (! [X2,X3] : (! [X4] : (((r1_lattice3(X1,X2,X4) | ~r1_lattice3(X0,X2,X3)) & (r2_lattice3(X1,X2,X4) | ~r2_lattice3(X0,X2,X3))) | X3 != X4 | ~m1_subset_1(X4,u1_struct_0(X1))) | ~m1_subset_1(X3,u1_struct_0(X0))) | g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) | ~l1_orders_2(X1)) | ~l1_orders_2(X0)) [flattening 42]
46. ! [X0,X1] : (! [X2,X3] : ((X1 = X3 & X0 = X2) | g1_orders_2(X0,X1) != g1_orders_2(X2,X3)) | ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0)))) [ennf transformation 19]
47. ? [X0] : (? [X1] : (((~v2_yellow_0(X1) & v2_yellow_0(X0)) | (~v1_yellow_0(X1) & v1_yellow_0(X0))) & g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) & l1_orders_2(X1)) & l1_orders_2(X0)) => (? [X1] : (((~v2_yellow_0(X1) & v2_yellow_0(sK0)) | (~v1_yellow_0(X1) & v1_yellow_0(sK0))) & g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) = g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) & l1_orders_2(X1)) & l1_orders_2(sK0)) [choice axiom]
48. ? [X1] : (((~v2_yellow_0(X1) & v2_yellow_0(sK0)) | (~v1_yellow_0(X1) & v1_yellow_0(sK0))) & g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) = g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) & l1_orders_2(X1)) => (((~v2_yellow_0(sK1) & v2_yellow_0(sK0)) | (~v1_yellow_0(sK1) & v1_yellow_0(sK0))) & g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) = g1_orders_2(u1_struct_0(sK1),u1_orders_2(sK1)) & l1_orders_2(sK1)) [choice axiom]
49. (((~v2_yellow_0(sK1) & v2_yellow_0(sK0)) | (~v1_yellow_0(sK1) & v1_yellow_0(sK0))) & g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) = g1_orders_2(u1_struct_0(sK1),u1_orders_2(sK1)) & l1_orders_2(sK1)) & l1_orders_2(sK0) [skolemisation 36,48,47]
50. ! [X0] : (((v1_yellow_0(X0) | ! [X1] : (~r1_lattice3(X0,u1_struct_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)))) & (? [X1] : (r1_lattice3(X0,u1_struct_0(X0),X1) & m1_subset_1(X1,u1_struct_0(X0))) | ~v1_yellow_0(X0))) | ~l1_orders_2(X0)) [nnf transformation 40]
51. ! [X0] : (((v1_yellow_0(X0) | ! [X1] : (~r1_lattice3(X0,u1_struct_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)))) & (? [X2] : (r1_lattice3(X0,u1_struct_0(X0),X2) & m1_subset_1(X2,u1_struct_0(X0))) | ~v1_yellow_0(X0))) | ~l1_orders_2(X0)) [rectify 50]
52. ! [X0] : (? [X2] : (r1_lattice3(X0,u1_struct_0(X0),X2) & m1_subset_1(X2,u1_struct_0(X0))) => (r1_lattice3(X0,u1_struct_0(X0),sK2(X0)) & m1_subset_1(sK2(X0),u1_struct_0(X0)))) [choice axiom]
53. ! [X0] : (((v1_yellow_0(X0) | ! [X1] : (~r1_lattice3(X0,u1_struct_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)))) & ((r1_lattice3(X0,u1_struct_0(X0),sK2(X0)) & m1_subset_1(sK2(X0),u1_struct_0(X0))) | ~v1_yellow_0(X0))) | ~l1_orders_2(X0)) [skolemisation 51,52]
54. ! [X0] : (((v2_yellow_0(X0) | ! [X1] : (~r2_lattice3(X0,u1_struct_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)))) & (? [X1] : (r2_lattice3(X0,u1_struct_0(X0),X1) & m1_subset_1(X1,u1_struct_0(X0))) | ~v2_yellow_0(X0))) | ~l1_orders_2(X0)) [nnf transformation 41]
55. ! [X0] : (((v2_yellow_0(X0) | ! [X1] : (~r2_lattice3(X0,u1_struct_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)))) & (? [X2] : (r2_lattice3(X0,u1_struct_0(X0),X2) & m1_subset_1(X2,u1_struct_0(X0))) | ~v2_yellow_0(X0))) | ~l1_orders_2(X0)) [rectify 54]
56. ! [X0] : (? [X2] : (r2_lattice3(X0,u1_struct_0(X0),X2) & m1_subset_1(X2,u1_struct_0(X0))) => (r2_lattice3(X0,u1_struct_0(X0),sK3(X0)) & m1_subset_1(sK3(X0),u1_struct_0(X0)))) [choice axiom]
57. ! [X0] : (((v2_yellow_0(X0) | ! [X1] : (~r2_lattice3(X0,u1_struct_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)))) & ((r2_lattice3(X0,u1_struct_0(X0),sK3(X0)) & m1_subset_1(sK3(X0),u1_struct_0(X0))) | ~v2_yellow_0(X0))) | ~l1_orders_2(X0)) [skolemisation 55,56]
64. l1_orders_2(sK0) [cnf transformation 49] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. l1_orders_2(sK1) [cnf transformation 49] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
66. g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) = g1_orders_2(u1_struct_0(sK1),u1_orders_2(sK1)) [cnf transformation 49] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
67. v2_yellow_0(sK0) | v1_yellow_0(sK0) [cnf transformation 49] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
68. v2_yellow_0(sK0) | ~v1_yellow_0(sK1) [cnf transformation 49] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
69. ~v2_yellow_0(sK1) | v1_yellow_0(sK0) [cnf transformation 49] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
70. ~v2_yellow_0(sK1) | ~v1_yellow_0(sK1) [cnf transformation 49] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
71. m1_subset_1(u1_orders_2(X0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0)))) | ~l1_orders_2(X0) [cnf transformation 37] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. m1_subset_1(sK2(X0),u1_struct_0(X0)) | ~v1_yellow_0(X0) | ~l1_orders_2(X0) [cnf transformation 53] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
74. r1_lattice3(X0,u1_struct_0(X0),sK2(X0)) | ~v1_yellow_0(X0) | ~l1_orders_2(X0) [cnf transformation 53] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
75. ~r1_lattice3(X0,u1_struct_0(X0),X1) | v1_yellow_0(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 53] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
76. m1_subset_1(sK3(X0),u1_struct_0(X0)) | ~v2_yellow_0(X0) | ~l1_orders_2(X0) [cnf transformation 57] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
77. r2_lattice3(X0,u1_struct_0(X0),sK3(X0)) | ~v2_yellow_0(X0) | ~l1_orders_2(X0) [cnf transformation 57] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
78. ~r2_lattice3(X0,u1_struct_0(X0),X1) | v2_yellow_0(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 57] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
79. r2_lattice3(X1,X2,X4) | ~r2_lattice3(X0,X2,X3) | X3 != X4 | ~m1_subset_1(X4,u1_struct_0(X1)) | ~m1_subset_1(X3,u1_struct_0(X0)) | g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) | ~l1_orders_2(X1) | ~l1_orders_2(X0) [cnf transformation 43] {a:0,w:34,wCS:13600,sine:2,thAx:0,allAx:1,thDist:-1}
80. r1_lattice3(X1,X2,X4) | ~r1_lattice3(X0,X2,X3) | X3 != X4 | ~m1_subset_1(X4,u1_struct_0(X1)) | ~m1_subset_1(X3,u1_struct_0(X0)) | g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) | ~l1_orders_2(X1) | ~l1_orders_2(X0) [cnf transformation 43] {a:0,w:34,wCS:13600,sine:2,thAx:0,allAx:1,thDist:-1}
86. g1_orders_2(X0,X1) != g1_orders_2(X2,X3) | X0 = X2 | ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) [cnf transformation 46] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
91. g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) | ~r1_lattice3(X0,X2,X4) | ~m1_subset_1(X4,u1_struct_0(X1)) | ~m1_subset_1(X4,u1_struct_0(X0)) | r1_lattice3(X1,X2,X4) | ~l1_orders_2(X1) | ~l1_orders_2(X0) [equality resolution 80] {a:0,w:31,wCS:12400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
92. g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)) | ~r2_lattice3(X0,X2,X4) | ~m1_subset_1(X4,u1_struct_0(X1)) | ~m1_subset_1(X4,u1_struct_0(X0)) | r2_lattice3(X1,X2,X4) | ~l1_orders_2(X1) | ~l1_orders_2(X0) [equality resolution 79] {a:0,w:31,wCS:12400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
94. 1 <=> v1_yellow_0(sK1) [avatar definition]
96. ~v1_yellow_0(sK1) <- (~1) [avatar component clause 94] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. 2 <=> v2_yellow_0(sK1) [avatar definition]
101. ~1 | ~2 [avatar split clause 70,98,94]
103. 3 <=> v1_yellow_0(sK0) [avatar definition]
105. v1_yellow_0(sK0) <- (3) [avatar component clause 103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
106. 3 | ~2 [avatar split clause 69,98,103]
108. 4 <=> v2_yellow_0(sK0) [avatar definition]
110. v2_yellow_0(sK0) <- (4) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
111. ~1 | 4 [avatar split clause 68,108,94]
112. 3 | 4 [avatar split clause 67,108,103]
114. 5 <=> g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) = g1_orders_2(u1_struct_0(sK1),u1_orders_2(sK1)) [avatar definition]
116. g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) = g1_orders_2(u1_struct_0(sK1),u1_orders_2(sK1)) <- (5) [avatar component clause 114] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
117. 5 [avatar split clause 66,114]
119. 6 <=> l1_orders_2(sK1) [avatar definition]
121. l1_orders_2(sK1) <- (6) [avatar component clause 119] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
122. 6 [avatar split clause 65,119]
124. 7 <=> l1_orders_2(sK0) [avatar definition]
126. l1_orders_2(sK0) <- (7) [avatar component clause 124] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
127. 7 [avatar split clause 64,124]
176. g1_orders_2(X0,X1) != g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) | u1_struct_0(sK1) = X0 | ~m1_subset_1(u1_orders_2(sK1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(sK1),u1_struct_0(sK1)))) <- (5) [superposition 86,116] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
182. 14 <=> m1_subset_1(u1_orders_2(sK1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(sK1),u1_struct_0(sK1)))) [avatar definition]
184. ~m1_subset_1(u1_orders_2(sK1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(sK1),u1_struct_0(sK1)))) <- (~14) [avatar component clause 182] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
186. 15 <=> ! [X1,X0] : (g1_orders_2(X0,X1) != g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) | u1_struct_0(sK1) = X0) [avatar definition]
187. g1_orders_2(X0,X1) != g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) | u1_struct_0(sK1) = X0 <- (15) [avatar component clause 186] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
188. ~14 | 15 | ~5 [avatar split clause 176,114,186,182]
202. ~l1_orders_2(sK1) <- (~14) [resolution 184,71] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
205. $false <- (6, ~14) [subsumption resolution 202,121] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
206. ~6 | 14 [avatar contradiction clause 205]
209. g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) | ~r1_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(sK1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | r1_lattice3(sK1,X1,X2) | ~l1_orders_2(sK1) | ~l1_orders_2(X0) <- (5) [superposition 91,116] {a:1,w:31,wCS:3100,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
216. g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) | ~r1_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(sK1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | r1_lattice3(sK1,X1,X2) | ~l1_orders_2(X0) <- (5, 6) [subsumption resolution 209,121] {a:1,w:29,wCS:2900,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
219. g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) | ~r2_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(sK1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | r2_lattice3(sK1,X1,X2) | ~l1_orders_2(sK1) | ~l1_orders_2(X0) <- (5) [superposition 92,116] {a:1,w:31,wCS:3100,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
226. g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) | ~r2_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(sK1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | r2_lattice3(sK1,X1,X2) | ~l1_orders_2(X0) <- (5, 6) [subsumption resolution 219,121] {a:1,w:29,wCS:2900,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
239. u1_struct_0(sK0) = u1_struct_0(sK1) <- (15) [equality resolution 187] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
250. 20 <=> u1_struct_0(sK0) = u1_struct_0(sK1) [avatar definition]
252. u1_struct_0(sK0) = u1_struct_0(sK1) <- (20) [avatar component clause 250] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
253. 20 | ~15 [avatar split clause 239,186,250]
270. ~r1_lattice3(sK1,u1_struct_0(sK0),X0) | v1_yellow_0(sK1) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK1) <- (20) [superposition 75,252] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
273. ~r2_lattice3(sK1,u1_struct_0(sK0),X1) | v2_yellow_0(sK1) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~l1_orders_2(sK1) <- (20) [superposition 78,252] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
291. ~r1_lattice3(sK1,u1_struct_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK1) <- (~1, 20) [subsumption resolution 270,96] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
292. ~r1_lattice3(sK1,u1_struct_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 6, 20) [subsumption resolution 291,121] {a:3,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
305. ~r2_lattice3(sK1,u1_struct_0(sK0),X1) | v2_yellow_0(sK1) | ~m1_subset_1(X1,u1_struct_0(sK0)) <- (6, 20) [subsumption resolution 273,121] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
307. 26 <=> ! [X1] : (~r2_lattice3(sK1,u1_struct_0(sK0),X1) | ~m1_subset_1(X1,u1_struct_0(sK0))) [avatar definition]
308. ~r2_lattice3(sK1,u1_struct_0(sK0),X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) <- (26) [avatar component clause 307] {a:3,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
309. 2 | 26 | ~6 | ~20 [avatar split clause 305,250,119,307,98]
361. g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) | ~m1_subset_1(X2,u1_struct_0(sK0)) | ~r1_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | r1_lattice3(sK1,X1,X2) | ~l1_orders_2(X0) <- (5, 6, 20) [forward demodulation 216,252] {a:1,w:29,wCS:2900,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
363. g1_orders_2(u1_struct_0(X0),u1_orders_2(X0)) != g1_orders_2(u1_struct_0(sK0),u1_orders_2(sK0)) | ~m1_subset_1(X2,u1_struct_0(sK0)) | ~r2_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | r2_lattice3(sK1,X1,X2) | ~l1_orders_2(X0) <- (5, 6, 20) [forward demodulation 226,252] {a:1,w:29,wCS:2900,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
379. ~m1_subset_1(X0,u1_struct_0(sK0)) | ~r1_lattice3(sK0,X1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | r1_lattice3(sK1,X1,X0) | ~l1_orders_2(sK0) <- (5, 6, 20) [equality resolution 361] {a:2,w:18,wCS:1800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
380. ~m1_subset_1(X0,u1_struct_0(sK0)) | ~r1_lattice3(sK0,X1,X0) | r1_lattice3(sK1,X1,X0) | ~l1_orders_2(sK0) <- (5, 6, 20) [duplicate literal removal 379] {a:2,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
383. ~r1_lattice3(sK0,X1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | r1_lattice3(sK1,X1,X0) <- (5, 6, 7, 20) [subsumption resolution 380,126] {a:2,w:12,wCS:1200,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
394. ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) | r1_lattice3(sK1,u1_struct_0(sK0),sK2(sK0)) | ~v1_yellow_0(sK0) | ~l1_orders_2(sK0) <- (5, 6, 7, 20) [resolution 383,74] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
395. ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) | r1_lattice3(sK1,u1_struct_0(sK0),sK2(sK0)) | ~l1_orders_2(sK0) <- (3, 5, 6, 7, 20) [subsumption resolution 394,105] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
396. ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) | r1_lattice3(sK1,u1_struct_0(sK0),sK2(sK0)) <- (3, 5, 6, 7, 20) [subsumption resolution 395,126] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
397. ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) <- (~1, 3, 5, 6, 7, 20) [subsumption resolution 396,292] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
399. 29 <=> m1_subset_1(sK2(sK0),u1_struct_0(sK0)) [avatar definition]
401. ~m1_subset_1(sK2(sK0),u1_struct_0(sK0)) <- (~29) [avatar component clause 399] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
402. ~29 | 1 | ~3 | ~5 | ~6 | ~7 | ~20 [avatar split clause 397,250,124,119,114,103,94,399]
407. ~m1_subset_1(X0,u1_struct_0(sK0)) | ~r2_lattice3(sK0,X1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | r2_lattice3(sK1,X1,X0) | ~l1_orders_2(sK0) <- (5, 6, 20) [equality resolution 363] {a:2,w:18,wCS:1800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
408. ~m1_subset_1(X0,u1_struct_0(sK0)) | ~r2_lattice3(sK0,X1,X0) | r2_lattice3(sK1,X1,X0) | ~l1_orders_2(sK0) <- (5, 6, 20) [duplicate literal removal 407] {a:2,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
411. ~r2_lattice3(sK0,X1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | r2_lattice3(sK1,X1,X0) <- (5, 6, 7, 20) [subsumption resolution 408,126] {a:2,w:12,wCS:1200,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
421. ~v1_yellow_0(sK0) | ~l1_orders_2(sK0) <- (~29) [resolution 401,73] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
426. ~v1_yellow_0(sK0) <- (7, ~29) [subsumption resolution 421,126] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
427. ~3 | ~7 | 29 [avatar split clause 426,399,124,103]
442. ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) | r2_lattice3(sK1,u1_struct_0(sK0),sK3(sK0)) | ~v2_yellow_0(sK0) | ~l1_orders_2(sK0) <- (5, 6, 7, 20) [resolution 411,77] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
443. ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) | r2_lattice3(sK1,u1_struct_0(sK0),sK3(sK0)) | ~l1_orders_2(sK0) <- (4, 5, 6, 7, 20) [subsumption resolution 442,110] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
444. ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) | r2_lattice3(sK1,u1_struct_0(sK0),sK3(sK0)) <- (4, 5, 6, 7, 20) [subsumption resolution 443,126] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
445. ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) <- (4, 5, 6, 7, 20, 26) [subsumption resolution 444,308] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
447. 30 <=> m1_subset_1(sK3(sK0),u1_struct_0(sK0)) [avatar definition]
449. ~m1_subset_1(sK3(sK0),u1_struct_0(sK0)) <- (~30) [avatar component clause 447] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
450. ~30 | ~4 | ~5 | ~6 | ~7 | ~20 | ~26 [avatar split clause 445,307,250,124,119,114,108,447]
464. ~v2_yellow_0(sK0) | ~l1_orders_2(sK0) <- (~30) [resolution 449,76] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
466. ~l1_orders_2(sK0) <- (4, ~30) [subsumption resolution 464,110] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
467. $false <- (4, 7, ~30) [subsumption resolution 466,126] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
468. ~4 | ~7 | 30 [avatar contradiction clause 467]
469. $false [avatar sat refutation 101,106,111,112,117,122,127,188,206,253,309,402,427,450,468] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:323,thDist:-323}
% SZS output end Proof for yellow_0__t13_yellow_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 280
% Active clauses: 66
% Passive clauses: 106
% Final active clauses: 58
% Final passive clauses: 29
% Main loop iterations started: 77
% 
% >>> Simplifying Inferences
% Duplicate literals: 25
% Fw subsumption resolutions: 50
% Fw demodulations: 24
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 16
% Equational tautologies: 4
% Forward subsumptions: 26
% Fw demodulations to eq. taut.: 1
% 
% >>> Generating Inferences
% Binary resolution: 24
% Unit resulting resolution: 4
% Forward superposition: 53
% Backward superposition: 21
% Equality resolution: 15
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 24
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 42
% SAT solver unit clauses: 6
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% superposition: 0.001 s ( own 0.0000 s ) 
% passive container maintenance: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

