i: [1,197,0,2,1,1,0,0] l24_xreal_0
a: [3,314,0,2,1,197]
s: 314
i: [1,196,0,2,1,1,0,0] l24_xreal_0
a: [3,319,0,2,1,196]
s: 319
i: [1,199,0,3,1,1,0,0] l24_xreal_0
a: [3,304,0,3,1,199]
s: 304
i: [1,198,0,3,1,1,0,0] l24_xreal_0
a: [3,309,0,3,1,198]
s: 309
i: [1,200,0,5,1,1,0,0] l24_xreal_0
a: [3,299,0,5,1,200]
s: 299
i: [1,201,0,2,1,0,0,255] fc1_numbers
a: [3,324,0,2,1,201]
s: 324
i: [1,202,0,2,1,0,0,255] spc1_boole
a: [3,329,0,2,1,202]
s: 329
i: [1,203,0,2,1,0,0,255] fc1_xboole_0
a: [3,334,0,2,1,203]
s: 334
i: [1,212,0,2,1,0,0,255] fc6_ordinal1
a: [3,374,0,2,1,212]
s: 374
i: [1,211,0,2,1,0,0,255] fc6_ordinal1
a: [3,379,0,2,1,211]
s: 379
i: [1,213,0,2,1,0,0,255] spc1_numerals
a: [3,399,0,2,1,213]
s: 399
i: [1,275,0,2,1,0,0,2] rc2_xboole_0
a: [3,411,0,2,1,275]
s: 411
i: [1,276,0,2,1,0,0,255] rc4_ordinal1
a: [3,416,0,2,1,276]
s: 416
i: [1,277,0,2,1,0,0,1] rc1_xreal_0
a: [3,421,0,2,1,277]
s: 421
i: [1,278,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,426,0,2,1,278]
s: 426
i: [1,279,0,2,1,0,0,255] rc1_ordinal1
a: [3,431,0,2,1,279]
s: 431
i: [1,280,0,2,1,0,0,2] rc1_xxreal_0
a: [3,436,0,2,1,280]
s: 436
i: [1,281,0,2,1,0,0,2] rc1_xboole_0
a: [3,441,0,2,1,281]
s: 441
i: [1,285,0,2,1,0,0,255] rc2_ordinal1
a: [3,446,0,2,1,285]
s: 446
i: [1,284,0,2,1,0,0,255] rc2_ordinal1
a: [3,451,0,2,1,284]
s: 451
i: [1,283,0,2,1,0,0,255] rc2_ordinal1
a: [3,456,0,2,1,283]
s: 456
i: [1,282,0,2,1,0,0,255] rc2_ordinal1
a: [3,461,0,2,1,282]
s: 461
i: [1,287,0,2,1,0,0,255] rc5_ordinal1
a: [3,466,0,2,1,287]
s: 466
i: [1,286,0,2,1,0,0,255] rc5_ordinal1
a: [3,471,0,2,1,286]
s: 471
i: [1,289,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,476,0,2,1,289]
s: 476
i: [1,288,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,481,0,2,1,288]
s: 481
i: [1,291,0,2,1,0,0,255] rc3_xxreal_0
a: [3,486,0,2,1,291]
s: 486
i: [1,290,0,2,1,0,0,255] rc3_xxreal_0
a: [3,491,0,2,1,290]
s: 491
i: [1,293,0,2,1,0,0,255] rc2_xxreal_0
a: [3,496,0,2,1,293]
s: 496
i: [1,292,0,2,1,0,0,255] rc2_xxreal_0
a: [3,501,0,2,1,292]
s: 501
i: [1,295,0,2,1,0,0,2] rc4_xxreal_0
a: [3,506,0,2,1,295]
s: 506
i: [1,294,0,2,1,0,0,2] rc4_xxreal_0
a: [3,511,0,2,1,294]
s: 511
i: [1,262,0,3,1,0,0,255] reflexivity_r1_tarski
s: 262
i: [1,204,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,339,0,3,1,204]
s: 339
i: [1,205,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,344,0,3,1,205]
s: 344
r: 334
d: [2,513,0,2,1,2,23,334,344]
a: [3,517,0,2,1,513]
s: 517
i: [1,206,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,349,0,3,1,206]
s: 349
i: [1,207,0,3,1,0,0,255] t1_numerals
a: [3,354,0,3,1,207]
s: 354
i: [1,216,0,3,1,0,0,255] spc1_numerals
a: [3,384,0,3,1,216]
s: 384
i: [1,215,0,3,1,0,0,255] spc1_numerals
a: [3,389,0,3,1,215]
s: 389
d: [2,519,0,3,1,2,22,354,344]
a: [3,523,0,3,1,519]
s: 523
d: [2,525,0,3,1,2,22,389,339]
a: [3,529,0,3,1,525]
s: 529
i: [1,218,0,4,2,0,0,255] cc2_xreal_0
s: 218
d: [2,531,1,2,1,1,40,218,416]
a: [3,536,1,2,1,531]
s: 536
i: [1,219,0,4,2,0,0,255] cc2_xcmplx_0
s: 219
i: [1,220,0,4,2,0,0,255] cc6_ordinal1
s: 220
d: [2,532,1,2,1,1,40,218,466]
a: [3,541,1,2,1,532]
s: 541
i: [1,221,0,4,2,0,0,255] cc2_xxreal_0
s: 221
d: [2,543,1,2,1,1,40,219,416]
a: [3,548,1,2,1,543]
s: 548
i: [1,222,0,4,2,0,0,1] cc3_xreal_0
s: 222
i: [1,223,0,4,2,0,0,1] cc4_xreal_0
s: 223
d: [2,579,1,2,1,1,40,222,319]
a: [3,587,1,2,1,579]
s: 587
i: [1,228,0,4,2,0,0,255] cc1_ordinal1
s: 228
d: [2,580,1,2,1,1,40,222,314]
a: [3,592,1,2,1,580]
s: 592
i: [1,229,0,4,2,0,0,255] cc1_ordinal1
s: 229
i: [1,231,0,4,2,0,0,255] cc9_ordinal1
s: 231
d: [2,600,1,2,1,1,40,223,319]
a: [3,608,1,2,1,600]
s: 608
i: [1,232,0,4,2,0,0,255] cc7_ordinal1
s: 232
d: [2,601,1,2,1,1,40,223,314]
a: [3,613,1,2,1,601]
s: 613
i: [1,233,0,4,2,0,0,255] cc3_ordinal1
s: 233
i: [1,261,0,4,1,0,0,255] existence_m1_subset_1
s: 261
d: [2,544,1,2,1,1,40,219,466]
a: [3,553,1,2,1,544]
s: 553
i: [1,208,0,4,1,0,0,255] dt_k5_numbers
a: [3,359,0,4,1,208]
s: 359
d: [2,555,1,2,1,1,40,220,416]
a: [3,560,1,2,1,555]
s: 560
i: [1,209,0,4,1,0,0,255] dt_k6_numbers
a: [3,364,0,4,1,209]
s: 364
i: [1,214,0,4,1,0,0,255] spc1_numerals
a: [3,394,0,4,1,214]
s: 394
d: [2,556,1,2,1,1,40,220,466]
a: [3,565,1,2,1,556]
s: 565
d: [2,702,0,4,1,2,22,359,339]
a: [3,706,0,4,1,702]
s: 706
d: [2,567,1,2,1,1,40,221,416]
a: [3,572,1,2,1,567]
s: 572
d: [2,720,0,4,1,2,22,364,339]
a: [3,724,0,4,1,720]
s: 724
d: [2,726,0,4,1,2,22,394,339]
a: [3,730,0,4,1,726]
s: 730
d: [2,568,1,2,1,1,40,221,466]
a: [3,577,1,2,1,568]
s: 577
i: [1,234,0,5,2,0,0,255] t6_boole
s: 234
d: [2,582,1,2,1,1,40,222,421]
a: [3,597,1,2,1,582]
s: 597
i: [1,235,0,5,2,0,0,255] cc8_ordinal1
s: 235
i: [1,236,0,5,2,0,0,255] cc1_xreal_0
s: 236
d: [2,603,1,2,1,1,40,223,421]
a: [3,619,1,2,1,603]
s: 619
i: [1,237,0,5,2,0,0,255] cc1_xcmplx_0
s: 237
d: [2,622,1,2,1,1,40,228,374]
a: [3,628,1,2,1,622]
s: 628
i: [1,258,0,5,2,0,0,255] t7_boole
s: 258
i: [1,210,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,369,0,5,1,210]
s: 369
d: [2,623,1,2,1,1,40,228,431]
a: [3,633,1,2,1,623]
s: 633
i: [1,272,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,406,0,5,2,272]
s: 406
d: [2,635,1,2,1,1,40,229,374]
a: [3,641,1,2,1,635]
s: 641
d: [2,744,0,5,2,1,22,234,344]
s: 744
i: [1,242,0,6,3,0,0,255] cc2_ordinal1
s: 242
d: [2,785,1,5,2,2,44,304,744]
s: 785
i: [1,245,0,6,3,0,0,255] cc5_xxreal_0
s: 245
d: [2,792,2,4,2,3,40,785,406]
d: [2,793,2,2,1,4,21,792,613]
a: [3,797,2,2,1,793]
s: 797
i: [1,247,0,6,3,0,0,255] cc5_xxreal_0
s: 247
i: [1,248,0,6,3,0,0,255] cc3_xxreal_0
s: 248
d: [2,786,1,5,2,2,44,309,744]
s: 786
i: [1,250,0,6,3,0,0,255] cc3_xxreal_0
s: 250
d: [2,821,2,4,2,3,40,786,406]
d: [2,822,2,2,1,4,21,821,608]
a: [3,826,2,2,1,822]
s: 826
i: [1,252,0,6,3,0,0,255] cc7_xxreal_0
s: 252
i: [1,253,0,6,3,0,0,255] cc7_xxreal_0
s: 253
d: [2,784,1,7,2,2,44,299,744]
s: 784
i: [1,254,0,6,2,0,0,255] t1_subset
s: 254
d: [2,636,1,2,1,1,40,229,431]
a: [3,646,1,2,1,636]
s: 646
i: [1,263,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 263
i: [1,226,0,7,2,0,0,1] t2_arithm
s: 226
d: [2,648,1,2,1,1,40,231,511]
a: [3,654,1,2,1,648]
s: 654
i: [1,227,0,7,2,0,0,255] t3_arithm
s: 227
d: [2,649,1,2,1,1,40,231,441]
a: [3,659,1,2,1,649]
s: 659
i: [1,230,0,7,3,0,0,255] cc5_ordinal1
s: 230
i: [1,256,0,7,2,0,0,255] t3_subset
s: 256
d: [2,650,1,2,1,1,40,231,517]
a: [3,664,1,2,1,650]
s: 664
i: [1,257,0,7,3,0,0,2] t8_boole
s: 257
d: [2,666,1,2,1,1,40,232,511]
a: [3,672,1,2,1,666]
s: 672
i: [1,217,0,8,3,0,0,255] cc10_ordinal1
s: 217
i: [1,239,0,8,4,0,0,255] cc4_xxreal_0
s: 239
r: 406
r: 797
a: [3,796,2,2,1,793]
s: 796
r: 796
a: [3,403,0,2,1,272]
s: 403
r: 245
r: 247
r: 248
r: 239
d: [2,985,1,0,0,2,40,403,506]
e: 985
d: [2,984,1,0,0,2,40,403,501]
e: 984
d: [2,983,1,0,0,2,40,403,491]
e: 983
d: [2,982,2,0,0,2,40,403,577]
e: 982
d: [2,981,1,0,0,2,40,403,436]
e: 981
d: [2,980,2,0,0,2,40,403,619]
e: 980
d: [2,979,2,0,0,2,40,403,572]
e: 979
d: [2,978,2,0,0,2,40,403,613]
e: 978
d: [2,977,2,0,0,2,40,403,608]
e: 977
d: [2,974,2,0,0,2,52,608,403]
e: 974
d: [2,975,1,0,0,2,52,319,403,223]
e: 975
d: [2,976,1,0,0,2,52,416,403,221]
e: 976
r: 403
s: 797
s: 406
s: 245
d: [2,877,2,6,3,1,40,239,608]
d: [2,886,2,4,2,2,21,877,826]
a: [3,890,2,2,1,886]
s: 890
s: 247
d: [2,878,2,6,3,1,40,239,613]
d: [2,896,2,4,2,2,21,878,797]
a: [3,900,2,2,1,896]
s: 900
s: 248
s: 239
d: [2,667,1,2,1,1,40,232,441]
a: [3,677,1,2,1,667]
s: 677
i: [1,241,0,8,4,0,0,255] cc6_xxreal_0
s: 241
d: [2,668,1,2,1,1,40,232,517]
a: [3,682,1,2,1,668]
s: 682
i: [1,243,0,8,4,0,0,255] cc8_xxreal_0
s: 243
i: [1,255,0,8,3,0,0,255] t2_subset
s: 255
d: [2,684,1,2,1,1,40,233,511]
a: [3,690,1,2,1,684]
s: 690
i: [1,269,0,8,3,0,0,1] fc6_xreal_0
s: 269
d: [2,685,1,2,1,1,40,233,441]
a: [3,695,1,2,1,685]
s: 695
i: [1,270,0,8,3,0,0,1] fc3_xcmplx_0
s: 270
i: [1,260,0,9,3,0,0,255] t5_subset
s: 260
d: [2,686,1,2,1,1,40,233,517]
a: [3,700,1,2,1,686]
s: 700
i: [1,259,0,10,3,0,0,255] t4_subset
s: 259
d: [2,746,1,2,1,1,40,235,529]
a: [3,752,1,2,1,746]
s: 752
i: [1,273,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 273
i: [1,271,0,11,3,0,0,1] commutativity_k3_xcmplx_0
s: 271
d: [2,1167,1,7,3,1,40,273,304]
d: [2,1173,1,5,2,2,21,1167,613]
a: [3,1177,1,3,1,1173]
s: 1177
i: [1,268,0,12,5,0,0,1] fc8_xcmplx_0
s: 268
d: [2,1168,1,7,3,1,40,273,309]
d: [2,1179,1,5,2,2,21,1168,608]
a: [3,1183,1,3,1,1179]
s: 1183
i: [1,224,0,13,5,0,0,1] l5_xreal_0
s: 224
i: [1,267,0,14,4,0,0,255] existence_m2_subset_1
s: 267
d: [2,1239,1,16,4,1,40,224,299]
a: [3,1247,1,4,1,1239]
s: 1247
d: [2,1261,2,0,0,3,52,319,314,1247,269]
e: 1261
d: [2,1262,2,4,2,1,40,1247,269]
d: [2,1264,2,2,1,2,21,1262,314]
d: [2,1265,2,0,0,3,21,1264,319]
e: 1265
r: 1247
i: [1,264,0,15,5,0,0,255] dt_m2_subset_1
s: 264
a: [3,1246,1,4,1,1239]
s: 1246
i: [1,265,0,15,5,0,0,255] redefinition_m2_subset_1
s: 265
i: [1,266,0,15,5,0,0,255] redefinition_m2_subset_1
s: 266
d: [2,1279,2,4,1,1,40,1246,223]
a: [3,1284,2,4,1,1279]
s: 1284
i: [1,274,0,17,4,0,0,1] spc7_arithm
s: 274
d: [2,1280,2,4,1,1,40,1246,222]
a: [3,1289,2,4,1,1280]
s: 1289
i: [1,225,0,19,6,0,0,1] l23_xreal_0
s: 225
a: [3,1251,1,5,1,1239]
s: 1251
d: [2,1300,3,12,3,1,40,1284,239]
a: [3,1304,3,4,1,1300]
s: 1304
d: [2,1433,4,0,0,5,52,826,587,797,592,1304,268]
e: 1433
d: [2,1434,4,8,4,1,40,1304,268]
d: [2,1439,4,6,3,2,21,1434,592]
d: [2,1440,4,4,2,3,21,1439,797]
d: [2,1441,4,2,1,4,21,1440,587]
d: [2,1442,4,0,0,5,21,1441,826]
e: 1442
r: 1304
d: [2,759,1,2,1,1,40,236,384]
a: [3,764,1,2,1,759]
s: 764
a: [3,1303,3,4,1,1300]
s: 1303
d: [2,771,1,2,1,1,40,237,384]
a: [3,776,1,2,1,771]
s: 776
d: [2,799,1,4,2,1,40,245,486]
d: [2,800,1,2,1,2,21,799,491]
a: [3,804,1,2,1,800]
s: 804
a: [3,1308,3,4,1,1300]
s: 1308
d: [2,806,1,4,2,1,40,247,486]
d: [2,807,1,2,1,2,21,806,491]
a: [3,811,1,2,1,807]
s: 811
d: [2,1223,2,5,2,2,44,1177,744]
s: 1223
d: [2,814,1,4,2,1,40,248,496]
d: [2,815,1,2,1,2,21,814,501]
a: [3,819,1,2,1,815]
s: 819
d: [2,832,1,4,2,1,40,230,384]
a: [3,844,1,2,1,832]
s: 844
d: [2,1234,2,5,2,2,44,1183,744]
s: 1234
d: [2,866,1,4,2,1,40,217,706]
a: [3,871,1,2,1,866]
s: 871
d: [2,1169,2,9,4,2,40,273,785]
d: [2,1185,2,7,3,3,21,1169,613]
s: 1185
d: [2,881,1,6,3,1,40,239,436]
a: [3,935,1,2,1,881]
s: 935
r: 935
a: [3,934,1,2,1,881]
s: 934
d: [2,1170,2,9,4,2,40,273,786]
d: [2,1186,2,7,3,3,21,1170,608]
s: 1186
a: [3,939,1,2,1,881]
s: 939
d: [2,1432,2,7,2,2,44,1251,744]
s: 1432
d: [2,747,1,3,1,0,40,235,261]
a: [3,757,1,3,1,747]
s: 757
d: [2,760,1,3,1,0,40,236,261]
a: [3,769,1,3,1,760]
s: 769
d: [2,708,2,2,1,1,40,560,229]
a: [3,713,2,2,1,708]
s: 713
d: [2,772,1,3,1,0,40,237,261]
a: [3,781,1,3,1,772]
s: 781
d: [2,709,2,2,1,1,40,560,228]
a: [3,718,2,2,1,709]
s: 718
d: [2,1068,1,5,2,1,40,255,523]
d: [2,1073,1,3,1,2,21,1068,379]
a: [3,1077,1,3,1,1073]
s: 1077
d: [2,1070,1,5,2,1,40,255,384]
d: [2,1088,1,3,1,2,21,1070,324]
a: [3,1092,1,3,1,1088]
s: 1092
d: [2,732,2,2,1,1,40,565,229]
a: [3,737,2,2,1,732]
s: 737
d: [2,1071,1,5,2,1,40,255,529]
d: [2,1094,1,3,1,2,21,1071,379]
a: [3,1098,1,3,1,1094]
s: 1098
d: [2,733,2,2,1,1,40,565,228]
a: [3,742,2,2,1,733]
s: 742
d: [2,1267,1,11,4,1,40,264,724]
d: [2,1271,1,7,3,2,21,1267,706]
d: [2,1272,1,5,2,3,21,1271,379]
d: [2,1273,1,3,1,4,21,1272,324]
a: [3,1277,1,3,1,1273]
s: 1277
d: [2,1369,1,17,6,1,44,225,369]
d: [2,1382,1,15,5,1,18,1369]
a: [3,1388,1,3,1,1382]
s: 1388
r: 1388
d: [2,852,2,2,1,1,40,672,219]
a: [3,858,2,2,1,852]
s: 858
a: [3,1387,1,3,1,1382]
s: 1387
d: [2,853,2,2,1,1,40,672,218]
a: [3,863,2,2,1,853]
s: 863
d: [2,1371,1,19,7,0,44,225,226]
d: [2,1398,1,17,6,0,21,1371,222]
a: [3,1402,1,3,1,1398]
s: 1402
r: 1402
a: [3,1401,1,3,1,1398]
s: 1401
d: [2,879,2,6,3,1,40,239,572]
a: [3,909,2,2,1,879]
s: 909
r: 909
d: [2,1069,1,7,2,1,40,255,706]
a: [3,1082,1,4,1,1069]
s: 1082
a: [3,908,2,2,1,879]
s: 908
a: [3,1085,1,3,1,1069]
s: 1085
d: [2,788,1,5,2,2,44,523,744]
s: 788
a: [3,913,2,2,1,879]
s: 913
d: [2,834,1,5,2,0,40,230,261]
s: 834
d: [2,880,2,6,3,1,40,239,619]
a: [3,922,2,2,1,880]
s: 922
r: 922
d: [2,789,1,6,2,2,44,724,744]
s: 789
d: [2,867,1,6,2,0,40,217,261]
s: 867
a: [3,921,2,2,1,880]
s: 921
d: [2,1072,1,6,2,0,40,255,261]
s: 1072
a: [3,926,2,2,1,880]
s: 926
d: [2,1153,1,6,2,1,40,259,706]
s: 1153
d: [2,783,1,7,3,1,44,744,744]
s: 783
d: [2,882,2,6,3,1,40,239,577]
d: [2,945,2,4,2,2,21,882,471]
a: [3,949,2,2,1,945]
s: 949
d: [2,865,1,7,3,0,40,217,256]
s: 865
d: [2,1026,2,2,1,1,40,677,221]
a: [3,1033,2,2,1,1026]
s: 1033
d: [2,1138,1,7,2,0,40,260,261]
s: 1138
d: [2,1187,1,11,4,0,44,271,227]
d: [2,1208,1,9,3,0,18,1187]
a: [3,1211,1,7,2,1208]
s: 1211
d: [2,1028,2,2,1,1,40,677,219]
a: [3,1039,2,2,1,1028]
s: 1039
d: [2,1188,1,11,4,0,44,271,226]
d: [2,1207,1,9,3,0,18,1188]
a: [3,1215,1,7,2,1207]
s: 1215
d: [2,1029,2,2,1,1,40,677,218]
a: [3,1044,2,2,1,1029]
s: 1044
d: [2,1296,1,11,4,1,40,266,706]
d: [2,1298,1,9,3,2,21,1296,379]
d: [2,1299,1,7,2,3,21,1298,324]
s: 1299
d: [2,848,1,8,3,0,40,256,230]
s: 848
d: [2,1046,2,2,1,1,40,682,221]
a: [3,1054,2,2,1,1046]
s: 1054
d: [2,1112,1,8,3,0,40,269,223]
s: 1112
d: [2,1048,2,2,1,1,40,682,219]
a: [3,1060,2,2,1,1048]
s: 1060
d: [2,1113,1,8,3,0,40,269,222]
s: 1113
d: [2,1136,1,8,3,0,40,260,256]
s: 1136
d: [2,1049,2,2,1,1,40,682,218]
a: [3,1065,2,2,1,1049]
s: 1065
d: [2,1154,1,8,2,0,40,259,261]
s: 1154
d: [2,1100,2,2,1,1,40,690,229]
a: [3,1105,2,2,1,1100]
s: 1105
d: [2,829,1,9,3,1,44,226,744]
s: 829
d: [2,1152,1,9,3,0,40,259,256]
s: 1152
d: [2,1101,2,2,1,1,40,690,228]
a: [3,1110,2,2,1,1101]
s: 1110
d: [2,1067,1,10,3,0,40,255,256]
s: 1067
d: [2,1118,2,2,1,1,40,695,229]
a: [3,1123,2,2,1,1118]
s: 1123
a: [3,1391,1,10,3,1382]
s: 1391
d: [2,1375,1,17,6,1,44,225,369]
d: [2,1381,1,15,5,1,18,1375]
a: [3,1414,1,10,3,1381]
s: 1414
d: [2,1119,2,2,1,1,40,695,228]
a: [3,1128,2,2,1,1119]
s: 1128
d: [2,1314,1,15,4,1,44,274,369]
d: [2,1340,1,13,3,1,18,1314]
a: [3,1343,1,11,2,1340]
s: 1343
d: [2,1139,2,2,1,1,40,700,229]
a: [3,1145,2,2,1,1139]
s: 1145
d: [2,1320,1,17,5,0,44,274,226]
d: [2,1354,1,13,4,0,21,1320,270]
a: [3,1357,1,11,3,1354]
s: 1357
d: [2,1327,1,18,6,0,44,270,274]
d: [2,1329,1,16,5,0,18,1327]
d: [2,1366,1,12,4,0,21,1329,270]
s: 1366
d: [2,1140,2,2,1,1,40,700,228]
a: [3,1150,2,2,1,1140]
s: 1150
a: [3,1405,1,12,4,1398]
s: 1405
d: [2,1155,2,2,1,1,40,752,221]
a: [3,1162,2,2,1,1155]
s: 1162
d: [2,1377,1,19,7,0,44,225,226]
d: [2,1421,1,17,6,0,21,1377,222]
a: [3,1424,1,12,4,1421]
s: 1424
d: [2,1269,1,21,7,0,40,264,267]
d: [2,1270,1,13,4,0,18,1269]
s: 1270
d: [2,1453,2,2,1,1,40,844,229]
a: [3,1458,2,2,1,1453]
s: 1458
d: [2,1293,1,21,7,0,40,265,267]
d: [2,1294,1,13,4,0,18,1293]
s: 1294
d: [2,1454,2,2,1,1,40,844,228]
a: [3,1463,2,2,1,1454]
s: 1463
d: [2,1315,1,17,5,0,44,274,227]
d: [2,1339,1,15,4,0,18,1315]
a: [3,1347,1,13,3,1339]
s: 1347
d: [2,1316,1,17,5,0,44,274,226]
d: [2,1338,1,15,4,0,18,1316]
a: [3,1351,1,13,3,1338]
s: 1351
d: [2,1924,2,16,6,1,44,1366,1357]
d: [2,1949,2,14,5,1,18,1924]
a: [3,1956,2,2,1,1949]
s: 1956
r: 1366
r: 1357
r: 1351
r: 1347
r: 1343
r: 1215
r: 1211
r: 829
r: 274
r: 271
r: 270
r: 268
r: 227
r: 226
d: [2,2150,3,0,0,2,40,1956,858]
e: 2150
d: [2,2149,3,0,0,2,40,1956,476]
e: 2149
d: [2,2148,3,0,0,2,40,1956,553]
e: 2148
d: [2,2147,3,0,0,2,40,1956,1039]
e: 2147
d: [2,2146,3,0,0,2,40,1956,426]
e: 2146
d: [2,2145,3,0,0,2,40,1956,597]
e: 2145
d: [2,2144,3,0,0,2,40,1956,548]
e: 2144
d: [2,2143,3,0,0,2,40,1956,781]
e: 2143
d: [2,2142,3,0,0,2,40,1956,592]
e: 2142
d: [2,2141,3,0,0,2,40,1956,587]
e: 2141
d: [2,2140,3,0,0,2,40,1956,776]
e: 2140
d: [2,2139,3,0,0,2,40,1956,1289]
e: 2139
d: [2,2137,3,0,0,2,40,1956,1060]
e: 2137
d: [2,2122,3,0,0,2,52,1060,1956]
e: 2122
d: [2,2123,3,0,0,2,52,1277,1956,237]
e: 2123
d: [2,2124,3,0,0,2,52,1065,1956,222]
e: 2124
d: [2,2125,3,0,0,2,52,682,1956,219]
e: 2125
d: [2,2126,3,0,0,2,52,1065,1065,1956,1113]
e: 2126
d: [2,2127,3,0,0,3,52,1065,1246,1956,1113]
e: 2127
d: [2,2128,3,0,0,3,52,1065,764,1956,1113]
e: 2128
d: [2,2129,3,0,0,3,52,1065,319,1956,1113]
e: 2129
d: [2,2130,3,0,0,3,52,1065,314,1956,1113]
e: 2130
d: [2,2131,3,0,0,3,52,1065,769,1956,1113]
e: 2131
d: [2,2132,3,0,0,3,52,1065,536,1956,1113]
e: 2132
d: [2,2133,3,0,0,3,52,1065,421,1956,1113]
e: 2133
d: [2,2134,3,0,0,3,52,1065,1044,1956,1113]
e: 2134
d: [2,2135,3,0,0,3,52,1065,541,1956,1113]
e: 2135
d: [2,2136,3,0,0,3,52,1065,863,1956,1113]
e: 2136
r: 1956
s: 226
d: [2,1491,3,16,6,2,40,1432,225]
d: [2,1493,3,13,5,3,21,1491,1177]
d: [2,1494,3,11,4,4,21,1493,314]
d: [2,1495,3,9,3,5,21,1494,319]
s: 1495
s: 227
s: 270
d: [2,1496,2,3,1,1,40,757,221]
a: [3,1503,2,3,1,1496]
s: 1503
s: 271
d: [2,2238,4,13,5,5,44,1495,271]
d: [2,2285,4,11,4,6,21,2238,592]
d: [2,2286,4,9,3,6,21,2285,222]
s: 2286
d: [2,2307,5,9,4,6,44,2286,226]
d: [2,2310,5,7,3,7,21,2307,517]
d: [2,2311,5,4,2,8,21,2310,1183]
d: [2,2312,5,2,1,9,21,2311,1065]
d: [2,2313,5,0,0,9,21,2312,592]
e: 2313
d: [2,745,1,2,1,1,40,235,523]
f: 2313,1442,1265,1239,1179,1173,1049,1046,985,822,793,745,601,600,580,579,519,513,295,272,207,205,203,196,197,198,199,200
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_0__l24_xreal_0
% SZS output start Proof for xreal_0__l24_xreal_0
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers)))) [input l24_xreal_0]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers)))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
13. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
27. ! [X0,X1] : ((v1_xcmplx_0(X1) & v1_xcmplx_0(X0)) => k3_xcmplx_0(X0,X1) = k3_xcmplx_0(X1,X0)) [input commutativity_k3_xcmplx_0]
28. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
41. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
44. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k3_xcmplx_0(X0,X1))) [input fc6_xreal_0]
46. ! [X0,X1] : ((v1_xcmplx_0(X1) & ~v1_xboole_0(X1) & v1_xcmplx_0(X0) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k3_xcmplx_0(X0,X1))) [input fc8_xcmplx_0]
47. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ! [X2] : (v1_xreal_0(X2) => ((r1_xxreal_0(k6_numbers,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(k3_xcmplx_0(X0,X2),k3_xcmplx_0(X1,X2)))))) [input l23_xreal_0]
48. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input l5_xreal_0]
60. ? [X0] : (v1_xxreal_0(X0) & v1_xboole_0(X0)) [input rc4_xxreal_0]
63. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
66. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => r1_xxreal_0(X0,X0)) [input reflexivity_r1_xxreal_0]
72. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
74. ! [X0] : (v1_xcmplx_0(X0) => k6_numbers = k3_xcmplx_0(X0,k6_numbers)) [input t2_arithm]
80. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
94. ? [X0] : (? [X1] : ((r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
95. ? [X0] : (? [X1] : (r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 94]
97. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
100. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
101. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 13]
102. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
103. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 48]
104. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 103]
105. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(k3_xcmplx_0(X0,X2),k3_xcmplx_0(X1,X2)) | (~r1_xxreal_0(k6_numbers,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 47]
106. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(k3_xcmplx_0(X0,X2),k3_xcmplx_0(X1,X2)) | ~r1_xxreal_0(k6_numbers,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 105]
107. ! [X0] : (k6_numbers = k3_xcmplx_0(X0,k6_numbers) | ~v1_xcmplx_0(X0)) [ennf transformation 74]
114. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 80]
115. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
148. ! [X0,X1] : (~v1_xboole_0(k3_xcmplx_0(X0,X1)) | (~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0))) [ennf transformation 46]
149. ! [X0,X1] : (~v1_xboole_0(k3_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0)) [flattening 148]
150. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 44]
151. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 150]
154. ! [X0,X1] : (k3_xcmplx_0(X0,X1) = k3_xcmplx_0(X1,X0) | (~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0))) [ennf transformation 27]
155. ! [X0,X1] : (k3_xcmplx_0(X0,X1) = k3_xcmplx_0(X1,X0) | ~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0)) [flattening 154]
156. ! [X0,X1] : (r1_xxreal_0(X0,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 66]
157. ! [X0,X1] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 156]
158. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 28]
159. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 158]
162. ? [X0] : (? [X1] : (r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (r1_xxreal_0(k3_xcmplx_0(sK0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
163. ? [X1] : (r1_xxreal_0(k3_xcmplx_0(sK0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) => (r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) & ~r1_xxreal_0(sK1,k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) [choice axiom]
164. (r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) & ~r1_xxreal_0(sK1,k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 95,163,162]
194. ? [X0] : (v1_xxreal_0(X0) & v1_xboole_0(X0)) => (v1_xxreal_0(sK16) & v1_xboole_0(sK16)) [choice axiom]
195. v1_xxreal_0(sK16) & v1_xboole_0(sK16) [skolemisation 60,194]
196. v1_xreal_0(sK0) [cnf transformation 164] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
197. v1_xreal_0(sK1) [cnf transformation 164] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
198. ~r1_xxreal_0(sK0,k6_numbers) [cnf transformation 164] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
199. ~r1_xxreal_0(sK1,k6_numbers) [cnf transformation 164] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
200. r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) [cnf transformation 164] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
203. v1_xboole_0(k1_xboole_0) [cnf transformation 41] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
205. k6_numbers = k1_xboole_0 [cnf transformation 63] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
207. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 72] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
218. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 97] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
221. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 100] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
222. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 101] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
223. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 102] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
224. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 104] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
225. r1_xxreal_0(k3_xcmplx_0(X0,X2),k3_xcmplx_0(X1,X2)) | ~r1_xxreal_0(k6_numbers,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 106] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
226. k6_numbers = k3_xcmplx_0(X0,k6_numbers) | ~v1_xcmplx_0(X0) [cnf transformation 107] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
234. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 114] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
235. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 115] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
268. ~v1_xboole_0(k3_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) [cnf transformation 149] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
269. v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 151] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
271. k3_xcmplx_0(X0,X1) = k3_xcmplx_0(X1,X0) | ~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0) [cnf transformation 155] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
272. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 157] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
273. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 159] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
295. v1_xxreal_0(sK16) [cnf transformation 195] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
297. 1 <=> r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) [avatar definition]
299. r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) <- (1) [avatar component clause 297] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
300. 1 [avatar split clause 200,297]
302. 2 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
304. ~r1_xxreal_0(sK1,k6_numbers) <- (~2) [avatar component clause 302] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
305. ~2 [avatar split clause 199,302]
307. 3 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
309. ~r1_xxreal_0(sK0,k6_numbers) <- (~3) [avatar component clause 307] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
310. ~3 [avatar split clause 198,307]
312. 4 <=> v1_xreal_0(sK1) [avatar definition]
314. v1_xreal_0(sK1) <- (4) [avatar component clause 312] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
315. 4 [avatar split clause 197,312]
317. 5 <=> v1_xreal_0(sK0) [avatar definition]
319. v1_xreal_0(sK0) <- (5) [avatar component clause 317] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
320. 5 [avatar split clause 196,317]
332. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
334. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 332] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
335. 8 [avatar split clause 203,332]
342. 10 <=> k6_numbers = k1_xboole_0 [avatar definition]
344. k6_numbers = k1_xboole_0 <- (10) [avatar component clause 342] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
345. 10 [avatar split clause 205,342]
352. 12 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
354. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (12) [avatar component clause 352] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
355. 12 [avatar split clause 207,352]
402. 22 <=> ! [X1] : ~v1_xxreal_0(X1) [avatar definition]
403. ~v1_xxreal_0(X1) <- (22) [avatar component clause 402] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
405. 23 <=> ! [X0] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0)) [avatar definition]
406. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) <- (23) [avatar component clause 405] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
407. 22 | 23 [avatar split clause 272,405,402]
504. 43 <=> v1_xxreal_0(sK16) [avatar definition]
506. v1_xxreal_0(sK16) <- (43) [avatar component clause 504] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
507. 43 [avatar split clause 295,504]
513. v1_xboole_0(k6_numbers) <- (8, 10) [backward demodulation 334,344] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
515. 45 <=> v1_xboole_0(k6_numbers) [avatar definition]
517. v1_xboole_0(k6_numbers) <- (45) [avatar component clause 515] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
518. 45 | ~8 | ~10 [avatar split clause 513,342,332,515]
519. m1_subset_1(k6_numbers,k4_ordinal1) <- (10, 12) [forward demodulation 354,344] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
521. 46 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
523. m1_subset_1(k6_numbers,k4_ordinal1) <- (46) [avatar component clause 521] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
524. 46 | ~10 | ~12 [avatar split clause 519,352,342,521]
579. v1_xcmplx_0(sK0) <- (5) [resolution 222,319] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
580. v1_xcmplx_0(sK1) <- (4) [resolution 222,314] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
585. 56 <=> v1_xcmplx_0(sK0) [avatar definition]
587. v1_xcmplx_0(sK0) <- (56) [avatar component clause 585] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
588. 56 | ~5 [avatar split clause 579,317,585]
590. 57 <=> v1_xcmplx_0(sK1) [avatar definition]
592. v1_xcmplx_0(sK1) <- (57) [avatar component clause 590] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
593. 57 | ~4 [avatar split clause 580,312,590]
600. v1_xxreal_0(sK0) <- (5) [resolution 223,319] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
601. v1_xxreal_0(sK1) <- (4) [resolution 223,314] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
606. 59 <=> v1_xxreal_0(sK0) [avatar definition]
608. v1_xxreal_0(sK0) <- (59) [avatar component clause 606] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
609. 59 | ~5 [avatar split clause 600,317,606]
611. 60 <=> v1_xxreal_0(sK1) [avatar definition]
613. v1_xxreal_0(sK1) <- (60) [avatar component clause 611] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
614. 60 | ~4 [avatar split clause 601,312,611]
680. 71 <=> v7_ordinal1(k6_numbers) [avatar definition]
682. v7_ordinal1(k6_numbers) <- (71) [avatar component clause 680] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
744. k6_numbers = X0 | ~v1_xboole_0(X0) <- (10) [forward demodulation 234,344] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
745. v7_ordinal1(k6_numbers) <- (46) [resolution 235,523] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
748. 71 | ~46 [avatar split clause 745,521,680]
785. ~r1_xxreal_0(sK1,X0) | ~v1_xboole_0(X0) <- (~2, 10) [superposition 304,744] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
786. ~r1_xxreal_0(sK0,X0) | ~v1_xboole_0(X0) <- (~3, 10) [superposition 309,744] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
792. ~v1_xboole_0(sK1) | ~v1_xxreal_0(sK1) <- (~2, 10, 23) [resolution 785,406] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
793. ~v1_xboole_0(sK1) <- (~2, 10, 23, 60) [subsumption resolution 792,613] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
795. 88 <=> v1_xboole_0(sK1) [avatar definition]
797. ~v1_xboole_0(sK1) <- (~88) [avatar component clause 795] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
798. ~88 | 2 | ~10 | ~23 | ~60 [avatar split clause 793,611,405,342,302,795]
821. ~v1_xboole_0(sK0) | ~v1_xxreal_0(sK0) <- (~3, 10, 23) [resolution 786,406] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
822. ~v1_xboole_0(sK0) <- (~3, 10, 23, 59) [subsumption resolution 821,608] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
824. 92 <=> v1_xboole_0(sK0) [avatar definition]
826. ~v1_xboole_0(sK0) <- (~92) [avatar component clause 824] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
827. ~92 | 3 | ~10 | ~23 | ~59 [avatar split clause 822,606,405,342,307,824]
985. $false <- (22, 43) [resolution 403,506] {a:1,w:0,sine:255,thAx:0,allAx:3,thDist:-3}
986. ~22 | ~43 [avatar contradiction clause 985]
1046. v1_xxreal_0(k6_numbers) <- (71) [resolution 682,221] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1049. v1_xreal_0(k6_numbers) <- (71) [resolution 682,218] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1052. 120 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1055. 120 | ~71 [avatar split clause 1046,680,1052]
1063. 122 <=> v1_xreal_0(k6_numbers) [avatar definition]
1065. v1_xreal_0(k6_numbers) <- (122) [avatar component clause 1063] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
1066. 122 | ~71 [avatar split clause 1049,680,1063]
1167. r1_xxreal_0(k6_numbers,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k6_numbers) <- (~2) [resolution 273,304] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1168. r1_xxreal_0(k6_numbers,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k6_numbers) <- (~3) [resolution 273,309] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1173. r1_xxreal_0(k6_numbers,sK1) | ~v1_xxreal_0(k6_numbers) <- (~2, 60) [subsumption resolution 1167,613] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1175. 135 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
1177. r1_xxreal_0(k6_numbers,sK1) <- (135) [avatar component clause 1175] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1178. ~120 | 135 | 2 | ~60 [avatar split clause 1173,611,302,1175,1052]
1179. r1_xxreal_0(k6_numbers,sK0) | ~v1_xxreal_0(k6_numbers) <- (~3, 59) [subsumption resolution 1168,608] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1181. 136 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
1183. r1_xxreal_0(k6_numbers,sK0) <- (136) [avatar component clause 1181] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1184. ~120 | 136 | 3 | ~59 [avatar split clause 1179,606,307,1181,1052]
1239. k6_numbers = k3_xcmplx_0(sK0,sK1) | ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k3_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (1) [resolution 224,299] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1245. 139 <=> v1_xreal_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
1247. ~v1_xreal_0(k3_xcmplx_0(sK0,sK1)) <- (~139) [avatar component clause 1245] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1249. 140 <=> r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK1)) [avatar definition]
1251. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK1)) <- (~140) [avatar component clause 1249] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1253. 141 <=> k6_numbers = k3_xcmplx_0(sK0,sK1) [avatar definition]
1256. ~122 | ~139 | ~140 | 141 | ~1 [avatar split clause 1239,297,1253,1249,1245,1063]
1262. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~139) [resolution 1247,269] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1264. ~v1_xreal_0(sK0) <- (4, ~139) [subsumption resolution 1262,314] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1265. $false <- (4, 5, ~139) [subsumption resolution 1264,319] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1266. ~4 | ~5 | 139 [avatar contradiction clause 1265]
1302. 146 <=> v1_xboole_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
1304. v1_xboole_0(k3_xcmplx_0(sK0,sK1)) <- (146) [avatar component clause 1302] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1432. ~r1_xxreal_0(X0,k3_xcmplx_0(sK0,sK1)) | ~v1_xboole_0(X0) <- (10, ~140) [superposition 1251,744] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1434. ~v1_xcmplx_0(sK1) | v1_xboole_0(sK1) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (146) [resolution 1304,268] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1439. v1_xboole_0(sK1) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (57, 146) [subsumption resolution 1434,592] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1440. ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (57, ~88, 146) [subsumption resolution 1439,797] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1441. v1_xboole_0(sK0) <- (56, 57, ~88, 146) [subsumption resolution 1440,587] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1442. $false <- (56, 57, ~88, ~92, 146) [subsumption resolution 1441,826] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1443. ~56 | ~57 | 88 | 92 | ~146 [avatar contradiction clause 1442]
1491. ~v1_xboole_0(k3_xcmplx_0(X1,sK1)) | ~r1_xxreal_0(k6_numbers,sK1) | ~r1_xxreal_0(X1,sK0) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) | ~v1_xreal_0(X1) <- (10, ~140) [resolution 1432,225] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1493. ~v1_xboole_0(k3_xcmplx_0(X1,sK1)) | ~r1_xxreal_0(X1,sK0) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) | ~v1_xreal_0(X1) <- (10, 135, ~140) [subsumption resolution 1491,1177] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1494. ~v1_xboole_0(k3_xcmplx_0(X1,sK1)) | ~r1_xxreal_0(X1,sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(X1) <- (4, 10, 135, ~140) [subsumption resolution 1493,314] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1495. ~v1_xboole_0(k3_xcmplx_0(X1,sK1)) | ~r1_xxreal_0(X1,sK0) | ~v1_xreal_0(X1) <- (4, 5, 10, 135, ~140) [subsumption resolution 1494,319] {a:3,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2238. ~v1_xboole_0(k3_xcmplx_0(sK1,X21)) | ~r1_xxreal_0(X21,sK0) | ~v1_xreal_0(X21) | ~v1_xcmplx_0(sK1) | ~v1_xcmplx_0(X21) <- (4, 5, 10, 135, ~140) [superposition 1495,271] {a:4,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2285. ~v1_xboole_0(k3_xcmplx_0(sK1,X21)) | ~r1_xxreal_0(X21,sK0) | ~v1_xreal_0(X21) | ~v1_xcmplx_0(X21) <- (4, 5, 10, 57, 135, ~140) [subsumption resolution 2238,592] {a:4,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2286. ~v1_xboole_0(k3_xcmplx_0(sK1,X21)) | ~r1_xxreal_0(X21,sK0) | ~v1_xreal_0(X21) <- (4, 5, 10, 57, 135, ~140) [subsumption resolution 2285,222] {a:4,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2307. ~v1_xboole_0(k6_numbers) | ~r1_xxreal_0(k6_numbers,sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xcmplx_0(sK1) <- (4, 5, 10, 57, 135, ~140) [superposition 2286,226] {a:5,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
2310. ~r1_xxreal_0(k6_numbers,sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xcmplx_0(sK1) <- (4, 5, 10, 45, 57, 135, ~140) [subsumption resolution 2307,517] {a:5,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
2311. ~v1_xreal_0(k6_numbers) | ~v1_xcmplx_0(sK1) <- (4, 5, 10, 45, 57, 135, 136, ~140) [subsumption resolution 2310,1183] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
2312. ~v1_xcmplx_0(sK1) <- (4, 5, 10, 45, 57, 122, 135, 136, ~140) [subsumption resolution 2311,1065] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
2313. $false <- (4, 5, 10, 45, 57, 122, 135, 136, ~140) [subsumption resolution 2312,592] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
2314. ~4 | ~5 | ~10 | ~45 | ~57 | ~122 | ~135 | ~136 | 140 [avatar contradiction clause 2313]
2315. k6_numbers != k3_xcmplx_0(sK0,sK1) | k6_numbers != k1_xboole_0 | v1_xboole_0(k3_xcmplx_0(sK0,sK1)) | ~v1_xboole_0(k1_xboole_0) [theory tautology sat conflict] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
2316. $false [avatar sat refutation 300,305,310,315,320,335,345,355,407,507,518,524,588,593,609,614,748,798,827,986,1055,1066,1178,1184,1256,1266,1443,2314,2315] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1443,thDist:-1443}
% SZS output end Proof for xreal_0__l24_xreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 81
% >>> Preprocessing
% Introduced skolems: 17
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 100
% Generated clauses: 1304
% Active clauses: 235
% Passive clauses: 420
% Final active clauses: 204
% Final passive clauses: 169
% Main loop iterations started: 249
% 
% >>> Simplifying Inferences
% Duplicate literals: 172
% Fw subsumption resolutions: 186
% Fw demodulations: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 60
% Equational tautologies: 3
% Forward subsumptions: 232
% Backward subsumptions: 18
% 
% >>> Generating Inferences
% Binary resolution: 298
% Unit resulting resolution: 20
% Factoring: 1
% Forward superposition: 188
% Backward superposition: 154
% Self superposition: 5
% 
% >>> AVATAR
% Split clauses: 48
% Split components: 115
% Sat splitting refutations: 45
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 262
% SAT solver unit clauses: 45
% SAT solver binary clauses: 143
% 
% Memory used [KB]: 6396
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% congruence closure: 0.001 s
% forward subsumption: 0.002 s
% forward subsumption resolution: 0.001 s
% backward subsumption: 0.001 s
% superposition: 0.001 s
% other: 0.013 s ( own 0.006 s ) 

