i: [1,167,0,2,1,1,0,0] t24_relset_1
a: [3,277,0,2,1,167]
s: 277
i: [1,166,0,2,1,1,0,0] t24_relset_1
a: [3,282,0,2,1,166]
s: 282
i: [1,170,0,8,2,1,0,0] t24_relset_1
a: [3,262,0,3,1,170]
s: 262
i: [1,169,0,3,1,1,0,0] t24_relset_1
a: [3,267,0,3,1,169]
s: 267
i: [1,172,0,13,3,1,0,0] t24_relset_1
a: [3,249,0,5,1,172]
s: 249
i: [1,171,0,10,2,1,0,0] t24_relset_1
a: [3,257,0,5,1,171]
s: 257
i: [1,168,0,6,1,1,0,0] t24_relset_1
a: [3,272,0,6,1,168]
s: 272
i: [1,173,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,287,0,2,1,173]
s: 287
i: [1,174,0,2,1,0,0,255] fc1_xboole_0
a: [3,292,0,2,1,174]
s: 292
i: [1,239,0,2,1,0,0,1] rc2_xboole_0
a: [3,302,0,2,1,239]
s: 302
i: [1,240,0,2,1,0,0,1] rc1_xboole_0
a: [3,307,0,2,1,240]
s: 307
i: [1,242,0,2,1,0,0,2] rc1_relat_1
a: [3,312,0,2,1,242]
s: 312
i: [1,241,0,2,1,0,0,2] rc1_relat_1
a: [3,317,0,2,1,241]
s: 317
i: [1,243,0,2,1,0,0,255] rc2_relat_1
a: [3,322,0,2,1,243]
s: 322
i: [1,176,0,3,1,0,0,255] fc2_xboole_0
s: 176
i: [1,177,0,3,1,0,0,1] fc1_subset_1
s: 177
i: [1,216,0,3,1,0,0,1] rc2_subset_1
s: 216
i: [1,219,0,3,1,0,0,255] reflexivity_r1_tarski
s: 219
i: [1,175,0,3,1,0,0,255] d2_xboole_0
a: [3,297,0,3,1,175]
s: 297
r: 287
i: [1,185,0,4,2,0,0,2] cc1_relat_1
s: 185
i: [1,214,0,4,1,0,0,1] existence_m1_subset_1
s: 214
d: [2,325,1,2,1,1,40,185,292]
a: [3,331,1,2,1,325]
s: 331
i: [1,218,0,4,1,0,0,255] rc3_subset_1
s: 218
d: [2,327,1,2,1,1,40,185,307]
a: [3,336,1,2,1,327]
s: 336
i: [1,220,0,4,1,0,0,255] fc3_xboole_0
s: 220
i: [1,221,0,4,1,0,0,1] fc6_relat_1
s: 221
d: [2,326,1,3,1,0,40,185,216]
s: 326
i: [1,235,0,4,1,0,0,255] rc3_relat_1
s: 235
i: [1,181,0,5,2,0,0,1] rc1_subset_1
s: 181
i: [1,186,0,5,2,0,0,255] t6_boole
s: 186
i: [1,187,0,5,2,0,0,2] fc10_relat_1
s: 187
i: [1,208,0,5,2,0,0,1] t7_boole
s: 208
i: [1,215,0,5,1,0,0,1] rc2_subset_1
s: 215
d: [2,342,1,2,1,1,40,208,257]
a: [3,346,1,2,1,342]
s: 346
i: [1,217,0,5,1,0,0,255] rc3_subset_1
s: 217
i: [1,222,0,5,1,0,0,255] fc5_relat_1
s: 222
i: [1,236,0,5,1,0,0,255] rc3_relat_1
s: 236
i: [1,237,0,5,1,0,0,255] rc3_relat_1
s: 237
d: [2,341,1,5,2,0,40,187,185]
s: 341
i: [1,183,0,6,2,0,0,255] rc4_subset_1
s: 183
i: [1,200,0,6,2,0,0,1] t1_subset
s: 200
r: 249
r: 257
a: [3,248,0,5,1,172]
s: 248
r: 248
s: 249
s: 257
a: [3,252,0,8,2,172]
s: 252
d: [2,369,1,0,0,3,52,262,257,252]
e: 369
d: [2,370,1,3,1,2,40,252,257]
d: [2,372,1,0,0,3,21,370,262]
e: 372
r: 249
r: 257
s: 248
d: [2,355,1,4,1,1,40,248,208]
a: [3,364,1,4,1,355]
s: 364
i: [1,226,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 226
d: [2,348,1,5,1,1,40,200,257]
a: [3,352,1,5,1,348]
s: 352
i: [1,180,0,7,2,0,0,1] rc1_subset_1
s: 180
i: [1,182,0,7,2,0,0,255] rc4_subset_1
s: 182
d: [2,354,1,5,1,1,40,248,200]
a: [3,359,1,5,1,354]
s: 359
i: [1,196,0,7,3,0,0,2] fc8_relat_1
s: 196
d: [2,378,1,5,1,1,40,226,248]
a: [3,382,1,5,1,378]
s: 382
i: [1,206,0,7,2,0,0,255] t3_subset
s: 206
i: [1,207,0,7,3,0,0,2] t8_boole
s: 207
i: [1,223,0,7,1,0,0,255] commutativity_k2_tarski
s: 223
i: [1,184,0,8,3,0,0,2] cc2_relat_1
s: 184
d: [2,388,1,6,2,1,40,184,272]
d: [2,394,1,2,1,1,21,388,221]
a: [3,398,1,2,1,394]
s: 398
i: [1,188,0,8,3,0,0,1] cc1_subset_1
s: 188
i: [1,201,0,8,3,0,0,1] t2_subset
s: 201
r: 346
a: [3,345,1,2,1,342]
s: 345
a: [3,256,0,5,1,171]
s: 256
d: [2,416,1,5,2,1,40,201,267]
d: [2,433,1,3,1,2,21,416,282]
a: [3,437,1,3,1,433]
s: 437
i: [1,227,0,8,3,0,0,1] fc10_subset_1
s: 227
i: [1,232,0,8,2,0,0,1] cc1_relset_1
s: 232
d: [2,418,1,5,2,1,40,201,262]
d: [2,439,1,3,1,2,21,418,277]
a: [3,443,1,3,1,439]
s: 443
i: [1,238,0,8,1,0,0,255] fc7_relat_1
s: 238
d: [2,452,2,3,1,1,40,437,226]
a: [3,458,2,3,1,452]
s: 458
i: [1,189,0,9,3,0,0,255] cc4_subset_1
s: 189
i: [1,190,0,9,4,0,0,255] cc8_relat_1
s: 190
d: [2,467,2,3,1,1,40,443,226]
a: [3,473,2,3,1,467]
s: 473
i: [1,193,0,9,4,0,0,2] cc7_relat_1
s: 193
d: [2,401,1,6,2,1,40,188,272]
d: [2,407,1,4,1,2,21,401,346]
a: [3,411,1,4,1,407]
s: 411
i: [1,210,0,9,3,0,0,1] t5_subset
s: 210
i: [1,233,0,9,2,0,0,255] cc2_relset_1
s: 233
d: [2,501,1,3,1,1,40,233,272]
a: [3,510,1,3,1,501]
s: 510
i: [1,234,0,9,2,0,0,255] cc2_relset_1
s: 234
d: [2,514,1,3,1,1,40,234,272]
a: [3,523,1,3,1,514]
s: 523
i: [1,244,0,9,2,0,0,2] d12_xtuple_0
s: 244
i: [1,199,0,10,1,0,0,255] d5_tarski
s: 199
d: [2,415,1,11,2,1,40,201,272]
d: [2,427,1,6,1,1,21,415,177]
a: [3,431,1,6,1,427]
s: 431
i: [1,209,0,10,3,0,0,1] t4_subset
s: 209
d: [2,534,2,6,1,1,40,431,226]
a: [3,540,2,6,1,534]
s: 540
i: [1,211,0,10,2,0,0,1] t87_zfmisc_1
s: 211
i: [1,212,0,10,2,0,0,1] t87_zfmisc_1
s: 212
d: [2,543,1,8,2,1,40,209,272]
s: 543
i: [1,224,0,10,3,0,0,1] cc3_relset_1
s: 224
d: [2,549,2,12,3,1,40,543,201]
d: [2,550,2,8,2,2,21,549,411]
s: 550
i: [1,225,0,10,3,0,0,1] cc4_relset_1
s: 225
i: [1,178,0,11,4,0,0,255] cc3_subset_1
s: 178
r: 550
r: 411
r: 256
r: 262
r: 345
r: 443
s: 346
a: [3,261,0,3,1,170]
s: 261
a: [3,442,1,3,1,439]
s: 442
s: 257
i: [1,179,0,11,4,0,0,255] cc2_subset_1
s: 179
s: 411
i: [1,229,0,11,3,0,0,1] redefinition_k1_relset_1
s: 229
r: 248
r: 252
r: 261
r: 359
r: 442
s: 262
s: 249
r: 249
r: 262
s: 261
s: 442
s: 248
d: [2,615,2,8,3,1,44,364,229]
d: [2,631,2,5,2,2,21,615,510]
d: [2,632,2,3,1,3,21,631,398]
a: [3,636,2,3,1,632]
s: 636
s: 252
i: [1,230,0,11,3,0,0,1] dt_k1_relset_1
s: 230
d: [2,599,1,4,1,1,40,257,244]
a: [3,606,1,4,1,599]
s: 606
i: [1,245,0,11,2,0,0,2] d12_xtuple_0
s: 245
d: [2,613,2,9,3,1,44,382,229]
d: [2,617,2,6,2,2,21,613,510]
d: [2,618,2,4,1,3,21,617,398]
a: [3,622,2,4,1,618]
s: 622
i: [1,228,0,12,4,0,0,255] cc6_relat_1
s: 228
i: [1,231,0,12,4,0,0,2] cc5_relat_1
s: 231
d: [2,614,2,9,3,1,44,359,229]
d: [2,624,2,6,2,2,21,614,510]
d: [2,625,2,4,1,3,21,624,398]
a: [3,629,2,4,1,625]
s: 629
i: [1,213,0,13,3,0,0,1] t87_zfmisc_1
s: 213
s: 359
i: [1,205,0,16,3,0,0,2] d12_xtuple_0
s: 205
i: [1,204,0,18,3,0,0,2] d12_xtuple_0
s: 204
d: [2,478,1,9,2,1,40,189,272]
a: [3,487,1,5,1,478]
s: 487
d: [2,600,1,5,1,1,40,257,226]
a: [3,611,1,5,1,600]
s: 611
d: [2,681,1,9,2,1,40,245,252]
d: [2,685,1,5,1,2,21,681,606]
a: [3,689,1,5,1,685]
s: 689
d: [2,693,1,12,3,1,40,228,272]
d: [2,699,1,8,2,1,21,693,221]
s: 699
d: [2,702,1,12,3,1,40,231,272]
d: [2,708,1,8,2,1,21,702,221]
s: 708
a: [3,590,2,8,2,549]
s: 590
d: [2,533,1,4,1,0,44,220,199]
s: 533
d: [2,728,3,8,2,1,40,590,226]
s: 728
d: [2,393,1,5,2,0,40,184,217]
s: 393
d: [2,406,1,5,2,0,40,188,217]
s: 406
d: [2,735,4,10,2,1,40,728,590]
d: [2,736,4,5,1,1,18,735]
a: [3,740,4,5,1,736]
s: 740
d: [2,422,1,8,2,0,40,201,215]
d: [2,447,1,5,1,0,21,422,177]
s: 447
d: [2,731,3,8,2,1,40,590,212]
s: 731
d: [2,746,4,0,0,3,52,442,257,731]
e: 746
d: [2,747,4,3,1,2,40,731,257]
d: [2,751,4,0,0,3,21,747,442]
e: 751
r: 257
r: 352
s: 256
a: [3,351,1,5,1,348]
s: 351
d: [2,732,3,8,2,1,40,590,211]
s: 732
d: [2,423,1,8,2,0,40,201,217]
d: [2,448,1,5,1,0,21,423,177]
s: 448
d: [2,754,4,7,2,1,40,732,245]
s: 754
d: [2,466,1,5,1,0,40,232,217]
s: 466
d: [2,579,1,9,3,0,40,178,217]
d: [2,586,1,5,2,0,21,579,218]
s: 586
d: [2,749,4,9,2,1,40,731,245]
s: 749
d: [2,391,1,6,2,0,40,184,214]
s: 391
d: [2,762,5,9,2,1,40,749,226]
s: 762
d: [2,404,1,6,2,0,40,188,214]
s: 404
d: [2,421,1,6,2,0,40,201,214]
s: 421
d: [2,769,5,8,2,1,40,421,754]
d: [2,770,5,5,1,2,21,769,636]
a: [3,774,5,5,1,770]
s: 774
d: [2,464,1,6,1,0,40,232,214]
s: 464
d: [2,776,6,5,1,1,40,774,226]
a: [3,782,6,5,1,776]
s: 782
d: [2,482,1,6,2,0,40,189,215]
s: 482
d: [2,489,1,10,3,0,40,190,237]
d: [2,490,1,6,2,0,21,489,235]
s: 490
d: [2,777,6,5,1,1,40,774,200]
a: [3,787,6,5,1,777]
s: 787
d: [2,491,1,10,3,0,40,193,236]
d: [2,492,1,6,2,0,21,491,235]
s: 492
d: [2,763,5,9,2,1,40,749,200]
s: 763
d: [2,792,6,0,0,3,52,606,689,763]
e: 792
d: [2,793,6,4,1,2,40,763,689]
d: [2,796,6,0,0,3,21,793,606]
e: 796
r: 346
r: 411
r: 763
r: 762
r: 754
r: 749
r: 732
r: 731
r: 728
r: 590
d: [2,375,1,4,1,1,40,248,208]
d: [2,644,1,9,3,1,44,249,229]
d: [2,645,1,6,2,2,21,644,510]
d: [2,646,1,4,1,3,21,645,398]
d: [2,651,1,9,3,1,44,248,229]
d: [2,653,1,6,2,2,21,651,510]
d: [2,654,1,4,1,3,21,653,398]
d: [2,656,3,2,1,1,40,636,187]
f: 656,796,685,654,646,632,599,549,501,401,394,375,168,171,172
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relset_1__t24_relset_1
% SZS output start Proof for relset_1__t24_relset_1
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ! [X3] : (m1_subset_1(X3,X0) => (r2_hidden(X3,k1_relset_1(X0,X2)) <=> ? [X4] : (r2_hidden(k4_tarski(X3,X4),X2) & m1_subset_1(X4,X1))))))) [input t24_relset_1]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ! [X3] : (m1_subset_1(X3,X0) => (r2_hidden(X3,k1_relset_1(X0,X2)) <=> ? [X4] : (r2_hidden(k4_tarski(X3,X4),X2) & m1_subset_1(X4,X1))))))) [negated conjecture 1]
6. ! [X0] : (v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_xboole_0(X1))) [input cc1_subset_1]
7. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
8. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [input cc2_relset_1]
21. ! [X0] : ! [X1] : (k9_xtuple_0(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : r2_hidden(k4_tarski(X2,X3),X0))) [input d12_xtuple_0]
35. ! [X0] : (v1_xboole_0(X0) => v1_xboole_0(k9_xtuple_0(X0))) [input fc10_relat_1]
42. ! [X0,X1] : v1_relat_1(k2_zfmisc_1(X0,X1)) [input fc6_relat_1]
54. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => k1_relset_1(X0,X1) = k9_xtuple_0(X1)) [input redefinition_k1_relset_1]
56. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
57. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
59. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
62. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
63. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [input t87_zfmisc_1]
66. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 56]
67. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 57]
68. ! [X0,X1] : (k9_xtuple_0(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : r2_hidden(k4_tarski(X2,X3),X0))) [flattening 21]
71. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 62]
72. ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 59]
74. ! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [flattening 63]
77. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [flattening 8]
82. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((r2_hidden(X3,k1_relset_1(X0,X2)) <~> ? [X4] : (r2_hidden(k4_tarski(X3,X4),X2) & m1_subset_1(X4,X1))) & m1_subset_1(X3,X0)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [ennf transformation 2]
89. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 7]
92. ! [X0] : (v1_xboole_0(k9_xtuple_0(X0)) | ~v1_xboole_0(X0)) [ennf transformation 35]
93. ! [X0] : (! [X1] : (v1_xboole_0(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_xboole_0(X0)) [ennf transformation 6]
105. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 66]
106. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 67]
107. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 106]
110. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 71]
111. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))) [ennf transformation 72]
112. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [flattening 111]
121. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 54]
122. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 121]
128. ! [X0,X1,X2] : ((v5_relat_1(X2,X1) & v4_relat_1(X2,X0)) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 77]
129. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((! [X4] : (~r2_hidden(k4_tarski(X3,X4),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k1_relset_1(X0,X2))) & (? [X4] : (r2_hidden(k4_tarski(X3,X4),X2) & m1_subset_1(X4,X1)) | r2_hidden(X3,k1_relset_1(X0,X2)))) & m1_subset_1(X3,X0)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [nnf transformation 82]
130. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X3,X4),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k1_relset_1(X0,X2))) & (? [X4] : (r2_hidden(k4_tarski(X3,X4),X2) & m1_subset_1(X4,X1)) | r2_hidden(X3,k1_relset_1(X0,X2))) & m1_subset_1(X3,X0)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [flattening 129]
131. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X3,X4),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k1_relset_1(X0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X3,X5),X2) & m1_subset_1(X5,X1)) | r2_hidden(X3,k1_relset_1(X0,X2))) & m1_subset_1(X3,X0)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [rectify 130]
132. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X3,X4),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k1_relset_1(X0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X3,X5),X2) & m1_subset_1(X5,X1)) | r2_hidden(X3,k1_relset_1(X0,X2))) & m1_subset_1(X3,X0)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X3,X4),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k1_relset_1(sK0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X3,X5),X2) & m1_subset_1(X5,X1)) | r2_hidden(X3,k1_relset_1(sK0,X2))) & m1_subset_1(X3,sK0)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,X1)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
133. ? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X3,X4),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k1_relset_1(sK0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X3,X5),X2) & m1_subset_1(X5,X1)) | r2_hidden(X3,k1_relset_1(sK0,X2))) & m1_subset_1(X3,sK0)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,X1)))) & ~v1_xboole_0(X1)) => (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X3,X4),X2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(X3,k1_relset_1(sK0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X3,X5),X2) & m1_subset_1(X5,sK1)) | r2_hidden(X3,k1_relset_1(sK0,X2))) & m1_subset_1(X3,sK0)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1)))) & ~v1_xboole_0(sK1)) [choice axiom]
134. ? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X3,X4),X2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(X3,k1_relset_1(sK0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X3,X5),X2) & m1_subset_1(X5,sK1)) | r2_hidden(X3,k1_relset_1(sK0,X2))) & m1_subset_1(X3,sK0)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1)))) => (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X3,X4),sK2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(X3,k1_relset_1(sK0,sK2))) & (? [X5] : (r2_hidden(k4_tarski(X3,X5),sK2) & m1_subset_1(X5,sK1)) | r2_hidden(X3,k1_relset_1(sK0,sK2))) & m1_subset_1(X3,sK0)) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1)))) [choice axiom]
135. ? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X3,X4),sK2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(X3,k1_relset_1(sK0,sK2))) & (? [X5] : (r2_hidden(k4_tarski(X3,X5),sK2) & m1_subset_1(X5,sK1)) | r2_hidden(X3,k1_relset_1(sK0,sK2))) & m1_subset_1(X3,sK0)) => ((! [X4] : (~r2_hidden(k4_tarski(sK3,X4),sK2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(sK3,k1_relset_1(sK0,sK2))) & (? [X5] : (r2_hidden(k4_tarski(sK3,X5),sK2) & m1_subset_1(X5,sK1)) | r2_hidden(sK3,k1_relset_1(sK0,sK2))) & m1_subset_1(sK3,sK0)) [choice axiom]
136. ? [X5] : (r2_hidden(k4_tarski(sK3,X5),sK2) & m1_subset_1(X5,sK1)) => (r2_hidden(k4_tarski(sK3,sK4),sK2) & m1_subset_1(sK4,sK1)) [choice axiom]
137. ((((! [X4] : (~r2_hidden(k4_tarski(sK3,X4),sK2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(sK3,k1_relset_1(sK0,sK2))) & ((r2_hidden(k4_tarski(sK3,sK4),sK2) & m1_subset_1(sK4,sK1)) | r2_hidden(sK3,k1_relset_1(sK0,sK2))) & m1_subset_1(sK3,sK0)) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1)))) & ~v1_xboole_0(sK1)) & ~v1_xboole_0(sK0) [skolemisation 131,136,135,134,133,132]
142. ! [X0,X1] : ((k9_xtuple_0(X0) = X1 | ? [X2] : ((! [X3] : ~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(X2,X1)) & (? [X3] : r2_hidden(k4_tarski(X2,X3),X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ! [X3] : ~r2_hidden(k4_tarski(X2,X3),X0)) & (? [X3] : r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(X2,X1))) | k9_xtuple_0(X0) != X1)) [nnf transformation 68]
143. ! [X0,X1] : ((k9_xtuple_0(X0) = X1 | ? [X2] : ((! [X3] : ~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(X2,X1)) & (? [X4] : r2_hidden(k4_tarski(X2,X4),X0) | r2_hidden(X2,X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : ~r2_hidden(k4_tarski(X5,X6),X0)) & (? [X7] : r2_hidden(k4_tarski(X5,X7),X0) | ~r2_hidden(X5,X1))) | k9_xtuple_0(X0) != X1)) [rectify 142]
144. ! [X1,X0] : (? [X2] : ((! [X3] : ~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(X2,X1)) & (? [X4] : r2_hidden(k4_tarski(X2,X4),X0) | r2_hidden(X2,X1))) => ((! [X3] : ~r2_hidden(k4_tarski(sK7(X0,X1),X3),X0) | ~r2_hidden(sK7(X0,X1),X1)) & (? [X4] : r2_hidden(k4_tarski(sK7(X0,X1),X4),X0) | r2_hidden(sK7(X0,X1),X1)))) [choice axiom]
145. ! [X1,X0] : (? [X4] : r2_hidden(k4_tarski(sK7(X0,X1),X4),X0) => r2_hidden(k4_tarski(sK7(X0,X1),sK8(X0,X1)),X0)) [choice axiom]
146. ! [X5,X0] : (? [X7] : r2_hidden(k4_tarski(X5,X7),X0) => r2_hidden(k4_tarski(X5,sK9(X0,X5)),X0)) [choice axiom]
147. ! [X0,X1] : ((k9_xtuple_0(X0) = X1 | ((! [X3] : ~r2_hidden(k4_tarski(sK7(X0,X1),X3),X0) | ~r2_hidden(sK7(X0,X1),X1)) & (r2_hidden(k4_tarski(sK7(X0,X1),sK8(X0,X1)),X0) | r2_hidden(sK7(X0,X1),X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : ~r2_hidden(k4_tarski(X5,X6),X0)) & (r2_hidden(k4_tarski(X5,sK9(X0,X5)),X0) | ~r2_hidden(X5,X1))) | k9_xtuple_0(X0) != X1)) [skolemisation 143,146,145,144]
148. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | (~r2_hidden(X1,X3) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [nnf transformation 74]
149. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [flattening 148]
168. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [cnf transformation 137] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
171. r2_hidden(k4_tarski(sK3,sK4),sK2) | r2_hidden(sK3,k1_relset_1(sK0,sK2)) [cnf transformation 137] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
172. ~r2_hidden(k4_tarski(sK3,X4),sK2) | ~m1_subset_1(X4,sK1) | ~r2_hidden(sK3,k1_relset_1(sK0,sK2)) [cnf transformation 137] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
184. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 89] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
187. v1_xboole_0(k9_xtuple_0(X0)) | ~v1_xboole_0(X0) [cnf transformation 92] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
188. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 93] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
200. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 105] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
201. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 107] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
202. r2_hidden(k4_tarski(X5,sK9(X0,X5)),X0) | ~r2_hidden(X5,X1) | k9_xtuple_0(X0) != X1 [cnf transformation 147] {a:0,w:14,wCS:5600,sine:2,thAx:0,allAx:4,thDist:-4}
203. r2_hidden(X5,X1) | ~r2_hidden(k4_tarski(X5,X6),X0) | k9_xtuple_0(X0) != X1 [cnf transformation 147] {a:0,w:12,wCS:4800,sine:2,thAx:0,allAx:4,thDist:-4}
208. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 110] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
209. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1) [cnf transformation 112] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
212. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X1,X3) [cnf transformation 149] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
221. v1_relat_1(k2_zfmisc_1(X0,X1)) [cnf transformation 42] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
229. k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 122] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
233. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v4_relat_1(X2,X0) [cnf transformation 128] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
244. ~r2_hidden(k4_tarski(X5,X6),X0) | r2_hidden(X5,k9_xtuple_0(X0)) [equality resolution 203] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:4,thDist:-4}
245. r2_hidden(k4_tarski(X5,sK9(X0,X5)),X0) | ~r2_hidden(X5,k9_xtuple_0(X0)) [equality resolution 202] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:4,thDist:-4}
247. 1 <=> r2_hidden(sK3,k1_relset_1(sK0,sK2)) [avatar definition]
248. r2_hidden(sK3,k1_relset_1(sK0,sK2)) <- (1) [avatar component clause 247] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
249. ~r2_hidden(sK3,k1_relset_1(sK0,sK2)) <- (~1) [avatar component clause 247] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
251. 2 <=> ! [X4] : (~r2_hidden(k4_tarski(sK3,X4),sK2) | ~m1_subset_1(X4,sK1)) [avatar definition]
252. ~r2_hidden(k4_tarski(sK3,X4),sK2) | ~m1_subset_1(X4,sK1) <- (2) [avatar component clause 251] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
253. ~1 | 2 [avatar split clause 172,251,247]
255. 3 <=> r2_hidden(k4_tarski(sK3,sK4),sK2) [avatar definition]
257. r2_hidden(k4_tarski(sK3,sK4),sK2) <- (3) [avatar component clause 255] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
258. 1 | 3 [avatar split clause 171,255,247]
270. 6 <=> m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [avatar definition]
272. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) <- (6) [avatar component clause 270] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
273. 6 [avatar split clause 168,270]
344. 19 <=> v1_xboole_0(sK2) [avatar definition]
362. 22 <=> v1_xboole_0(k1_relset_1(sK0,sK2)) [avatar definition]
364. ~v1_xboole_0(k1_relset_1(sK0,sK2)) <- (~22) [avatar component clause 362] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
375. ~v1_xboole_0(k1_relset_1(sK0,sK2)) <- (1) [resolution 248,208] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
377. ~22 | ~1 [avatar split clause 375,247,362]
388. v1_relat_1(sK2) | ~v1_relat_1(k2_zfmisc_1(sK0,sK1)) <- (6) [resolution 184,272] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
394. v1_relat_1(sK2) <- (6) [subsumption resolution 388,221] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
396. 24 <=> v1_relat_1(sK2) [avatar definition]
398. v1_relat_1(sK2) <- (24) [avatar component clause 396] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
399. 24 | ~6 [avatar split clause 394,270,396]
401. v1_xboole_0(sK2) | ~v1_xboole_0(k2_zfmisc_1(sK0,sK1)) <- (6) [resolution 188,272] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
409. 25 <=> v1_xboole_0(k2_zfmisc_1(sK0,sK1)) [avatar definition]
450. ~25 | 19 | ~6 [avatar split clause 401,270,344,409]
501. v4_relat_1(sK2,sK0) <- (6) [resolution 233,272] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
508. 32 <=> v4_relat_1(sK2,sK0) [avatar definition]
510. v4_relat_1(sK2,sK0) <- (32) [avatar component clause 508] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
511. 32 | ~6 [avatar split clause 501,270,508]
543. m1_subset_1(X3,k2_zfmisc_1(sK0,sK1)) | ~r2_hidden(X3,sK2) <- (6) [resolution 209,272] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
549. ~r2_hidden(X0,sK2) | v1_xboole_0(k2_zfmisc_1(sK0,sK1)) | r2_hidden(X0,k2_zfmisc_1(sK0,sK1)) <- (6) [resolution 543,201] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
589. 35 <=> ! [X0] : (~r2_hidden(X0,sK2) | r2_hidden(X0,k2_zfmisc_1(sK0,sK1))) [avatar definition]
590. r2_hidden(X0,k2_zfmisc_1(sK0,sK1)) | ~r2_hidden(X0,sK2) <- (35) [avatar component clause 589] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
591. 25 | 35 | ~6 [avatar split clause 549,270,589,409]
599. r2_hidden(sK3,k9_xtuple_0(sK2)) <- (3) [resolution 257,244] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
604. 37 <=> r2_hidden(sK3,k9_xtuple_0(sK2)) [avatar definition]
606. r2_hidden(sK3,k9_xtuple_0(sK2)) <- (37) [avatar component clause 604] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
607. 37 | ~3 [avatar split clause 599,255,604]
615. ~v1_xboole_0(k9_xtuple_0(sK2)) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~22) [superposition 364,229] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
631. ~v1_xboole_0(k9_xtuple_0(sK2)) | ~v1_relat_1(sK2) <- (~22, 32) [subsumption resolution 615,510] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
632. ~v1_xboole_0(k9_xtuple_0(sK2)) <- (~22, 24, 32) [subsumption resolution 631,398] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
634. 41 <=> v1_xboole_0(k9_xtuple_0(sK2)) [avatar definition]
636. ~v1_xboole_0(k9_xtuple_0(sK2)) <- (~41) [avatar component clause 634] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
637. ~41 | 22 | ~24 | ~32 [avatar split clause 632,508,396,362,634]
644. ~r2_hidden(sK3,k9_xtuple_0(sK2)) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1) [superposition 249,229] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
645. ~r2_hidden(sK3,k9_xtuple_0(sK2)) | ~v1_relat_1(sK2) <- (~1, 32) [subsumption resolution 644,510] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
646. ~r2_hidden(sK3,k9_xtuple_0(sK2)) <- (~1, 24, 32) [subsumption resolution 645,398] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
647. ~37 | 1 | ~24 | ~32 [avatar split clause 646,508,396,247,604]
651. r2_hidden(sK3,k9_xtuple_0(sK2)) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (1) [superposition 248,229] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
653. r2_hidden(sK3,k9_xtuple_0(sK2)) | ~v1_relat_1(sK2) <- (1, 32) [subsumption resolution 651,510] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
654. r2_hidden(sK3,k9_xtuple_0(sK2)) <- (1, 24, 32) [subsumption resolution 653,398] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
655. 37 | ~1 | ~24 | ~32 [avatar split clause 654,508,396,247,604]
656. ~v1_xboole_0(sK2) <- (~41) [resolution 636,187] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
681. ~r2_hidden(sK3,k9_xtuple_0(sK2)) | ~m1_subset_1(sK9(sK2,sK3),sK1) <- (2) [resolution 245,252] {a:1,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
685. ~m1_subset_1(sK9(sK2,sK3),sK1) <- (2, 37) [subsumption resolution 681,606] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
687. 42 <=> m1_subset_1(sK9(sK2,sK3),sK1) [avatar definition]
689. ~m1_subset_1(sK9(sK2,sK3),sK1) <- (~42) [avatar component clause 687] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
690. ~42 | ~2 | ~37 [avatar split clause 685,604,251,687]
731. ~r2_hidden(k4_tarski(X3,X4),sK2) | r2_hidden(X4,sK1) <- (35) [resolution 590,212] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
749. r2_hidden(sK9(sK2,X1),sK1) | ~r2_hidden(X1,k9_xtuple_0(sK2)) <- (35) [resolution 731,245] {a:4,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
763. m1_subset_1(sK9(sK2,X1),sK1) | ~r2_hidden(X1,k9_xtuple_0(sK2)) <- (35) [resolution 749,200] {a:5,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
793. ~r2_hidden(sK3,k9_xtuple_0(sK2)) <- (35, ~42) [resolution 763,689] {a:6,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:24,thDist:-24}
796. $false <- (35, 37, ~42) [subsumption resolution 793,606] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:24,thDist:-24}
797. ~35 | ~37 | 42 [avatar contradiction clause 796]
798. ~19 | 41 [avatar split clause 656,634,344]
799. $false [avatar sat refutation 253,258,273,377,399,450,511,591,607,637,647,655,690,797,798] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:716,thDist:-716}
% SZS output end Proof for relset_1__t24_relset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 63
% >>> Preprocessing
% Introduced skolems: 18
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 78
% Generated clauses: 467
% Active clauses: 148
% Passive clauses: 270
% Final active clauses: 116
% Final passive clauses: 107
% Main loop iterations started: 149
% 
% >>> Simplifying Inferences
% Duplicate literals: 6
% Fw subsumption resolutions: 41
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 20
% Forward subsumptions: 61
% 
% >>> Generating Inferences
% Binary resolution: 230
% Unit resulting resolution: 3
% Forward superposition: 10
% Backward superposition: 11
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 10
% Split components: 21
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 71
% SAT solver unit clauses: 12
% SAT solver binary clauses: 30
% 
% Memory used [KB]: 5373
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% term sharing: 0.001 s
% splitting model update: 0.001 s
% resolution: 0.001 s ( own 0.0000 s ) 
% unit resulting resolution: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

