i: [1,47,0,2,1,1,0,0] t68_xboole_1
a: [3,86,0,2,1,47]
s: 86
i: [1,50,0,3,1,1,0,0] t68_xboole_1
a: [3,71,0,3,1,50]
s: 71
i: [1,49,0,3,1,1,0,0] t68_xboole_1
a: [3,76,0,3,1,49]
s: 76
i: [1,48,0,3,1,1,0,0] t68_xboole_1
a: [3,81,0,3,1,48]
s: 81
i: [1,51,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,91,0,2,1,51]
s: 91
i: [1,52,0,2,1,0,0,255] fc1_xboole_0
a: [3,96,0,2,1,52]
s: 96
i: [1,66,0,2,1,0,0,1] rc2_xboole_0
a: [3,106,0,2,1,66]
s: 106
i: [1,67,0,2,1,0,0,1] rc1_xboole_0
a: [3,111,0,2,1,67]
s: 111
i: [1,63,0,3,1,0,0,1] reflexivity_r1_tarski
s: 63
i: [1,53,0,3,1,0,0,255] d2_xboole_0
a: [3,101,0,3,1,53]
s: 101
r: 91
i: [1,54,0,5,2,0,0,255] t6_boole
s: 54
i: [1,55,0,5,2,0,0,2] d1_xboole_0
s: 55
i: [1,62,0,5,2,0,0,2] t7_boole
s: 62
i: [1,56,0,6,2,0,0,2] d1_xboole_0
s: 56
i: [1,64,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 64
i: [1,65,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 65
d: [2,118,1,3,1,1,40,65,71]
a: [3,122,1,3,1,118]
s: 122
i: [1,61,0,7,3,0,0,255] t8_boole
s: 61
d: [2,117,1,6,2,0,40,64,56]
s: 117
i: [1,57,0,8,2,0,0,1] t3_xboole_0
s: 57
i: [1,58,0,8,2,0,0,1] t3_xboole_0
s: 58
d: [2,126,1,5,2,0,40,57,55]
s: 126
i: [1,59,0,9,3,0,0,1] t3_xboole_0
s: 59
d: [2,128,1,5,2,0,40,58,55]
s: 128
i: [1,60,0,9,3,0,0,1] d3_tarski
s: 60
d: [2,114,1,7,3,0,44,54,54]
s: 114
d: [2,129,2,5,2,0,40,126,65]
s: 129
d: [2,125,1,8,2,0,40,57,64]
s: 125
d: [2,127,1,8,2,0,40,58,64]
s: 127
d: [2,130,1,9,3,0,40,59,56]
s: 130
d: [2,134,1,9,3,0,40,60,56]
s: 134
d: [2,137,2,7,3,0,40,130,56]
d: [2,138,2,5,2,0,18,137]
s: 138
d: [2,131,1,11,3,0,40,59,57]
s: 131
d: [2,150,2,9,3,0,40,131,57]
d: [2,153,2,6,2,0,18,150]
s: 153
d: [2,132,1,11,3,0,40,59,58]
s: 132
d: [2,135,1,11,3,0,40,60,57]
s: 135
d: [2,157,2,9,3,0,40,132,58]
d: [2,158,2,6,2,0,18,157]
s: 158
d: [2,136,1,11,3,0,40,60,58]
s: 136
d: [2,143,2,7,3,0,40,134,55]
s: 143
d: [2,139,2,10,4,0,40,134,130]
d: [2,147,2,8,3,0,18,139]
s: 147
d: [2,165,2,8,3,0,40,135,55]
s: 165
d: [2,184,3,4,2,1,40,143,76]
d: [2,187,3,2,1,2,21,184,86]
a: [3,191,3,2,1,187]
s: 191
d: [2,178,2,8,3,0,40,136,55]
s: 178
d: [2,185,3,4,2,1,40,143,81]
d: [2,193,3,2,1,2,21,185,86]
a: [3,197,3,2,1,193]
s: 197
d: [2,142,2,9,3,0,40,134,64]
s: 142
d: [2,161,2,12,4,0,40,135,131]
d: [2,169,2,9,3,0,18,161]
s: 169
d: [2,199,3,5,2,1,40,147,71]
a: [3,206,3,3,1,199]
s: 206
d: [2,160,2,12,4,0,40,135,132]
d: [2,170,2,9,3,0,18,160]
s: 170
d: [2,200,3,5,2,1,40,147,122]
a: [3,211,3,3,1,200]
s: 211
d: [2,174,2,12,4,0,40,136,131]
d: [2,182,2,9,3,0,18,174]
s: 182
d: [2,173,2,12,4,0,40,136,132]
d: [2,183,2,9,3,0,18,173]
s: 183
d: [2,246,3,6,2,1,40,170,71]
s: 246
d: [2,144,2,10,3,0,40,134,117]
s: 144
d: [2,258,4,3,1,2,40,246,81]
a: [3,263,4,3,1,258]
s: 263
d: [2,164,2,11,3,0,40,135,64]
s: 164
d: [2,177,2,11,3,0,40,136,64]
s: 177
d: [2,265,5,6,2,1,40,263,183]
a: [3,274,5,3,1,265]
s: 274
d: [2,140,2,12,4,0,40,134,60]
s: 140
d: [2,270,5,3,1,1,40,263,65]
a: [3,285,5,3,1,270]
s: 285
d: [2,311,6,0,0,3,52,86,76,285,147]
e: 311
d: [2,316,6,5,2,1,40,285,147]
d: [2,321,6,3,1,2,21,316,86]
d: [2,322,6,0,0,3,21,321,76]
e: 322
f: 322,270,258,47,48,49,50
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t68_xboole_1
% SZS output start Proof for xboole_1__t68_xboole_1
1. ! [X0] : ! [X1] : ! [X2] : (~v1_xboole_0(X2) => ~(r1_xboole_0(X0,X1) & r1_tarski(X2,X1) & r1_tarski(X2,X0))) [input t68_xboole_1]
2. ~! [X0] : ! [X1] : ! [X2] : (~v1_xboole_0(X2) => ~(r1_xboole_0(X0,X1) & r1_tarski(X2,X1) & r1_tarski(X2,X0))) [negated conjecture 1]
4. ! [X0] : (v1_xboole_0(X0) <=> ! [X1] : ~r2_hidden(X1,X0)) [input d1_xboole_0]
6. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
13. ! [X0,X1] : (r1_xboole_0(X0,X1) => r1_xboole_0(X1,X0)) [input symmetry_r1_xboole_0]
14. ! [X0] : ! [X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X2] : ~(r2_hidden(X2,X1) & r2_hidden(X2,X0)) & ~r1_xboole_0(X0,X1))) [input t3_xboole_0]
18. ~! [X0,X1,X2] : (~v1_xboole_0(X2) => ~(r1_xboole_0(X0,X1) & r1_tarski(X2,X1) & r1_tarski(X2,X0))) [flattening 2]
19. ! [X0] : ! [X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X3] : ~(r2_hidden(X3,X1) & r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) [rectify 14]
20. ! [X0,X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X3] : ~(r2_hidden(X3,X1) & r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) [flattening 19]
21. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 6]
25. ! [X0,X1] : (r1_tarski(X0,X1) => ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [unused predicate definition removal 21]
26. ? [X0,X1,X2] : ((r1_xboole_0(X0,X1) & r1_tarski(X2,X1) & r1_tarski(X2,X0)) & ~v1_xboole_0(X2)) [ennf transformation 18]
27. ? [X0,X1,X2] : (r1_xboole_0(X0,X1) & r1_tarski(X2,X1) & r1_tarski(X2,X0) & ~v1_xboole_0(X2)) [flattening 26]
29. ! [X0,X1] : ((~r1_xboole_0(X0,X1) | ! [X2] : (~r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) & (? [X3] : (r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r1_xboole_0(X0,X1))) [ennf transformation 20]
30. ! [X0,X1] : (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1)) [ennf transformation 25]
34. ! [X0,X1] : (r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 13]
35. ? [X0,X1,X2] : (r1_xboole_0(X0,X1) & r1_tarski(X2,X1) & r1_tarski(X2,X0) & ~v1_xboole_0(X2)) => (r1_xboole_0(sK0,sK1) & r1_tarski(sK2,sK1) & r1_tarski(sK2,sK0) & ~v1_xboole_0(sK2)) [choice axiom]
36. r1_xboole_0(sK0,sK1) & r1_tarski(sK2,sK1) & r1_tarski(sK2,sK0) & ~v1_xboole_0(sK2) [skolemisation 27,35]
37. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X1] : ~r2_hidden(X1,X0) | ~v1_xboole_0(X0))) [nnf transformation 4]
38. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [rectify 37]
39. ! [X0] : (? [X1] : r2_hidden(X1,X0) => r2_hidden(sK3(X0),X0)) [choice axiom]
40. ! [X0] : ((v1_xboole_0(X0) | r2_hidden(sK3(X0),X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [skolemisation 38,39]
41. ! [X1,X0] : (? [X3] : (r2_hidden(X3,X1) & r2_hidden(X3,X0)) => (r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK4(X0,X1),X0))) [choice axiom]
42. ! [X0,X1] : ((~r1_xboole_0(X0,X1) | ! [X2] : (~r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) & ((r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK4(X0,X1),X0)) | r1_xboole_0(X0,X1))) [skolemisation 29,41]
47. ~v1_xboole_0(sK2) [cnf transformation 36] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. r1_tarski(sK2,sK0) [cnf transformation 36] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. r1_tarski(sK2,sK1) [cnf transformation 36] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. r1_xboole_0(sK0,sK1) [cnf transformation 36] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
56. r2_hidden(sK3(X0),X0) | v1_xboole_0(X0) [cnf transformation 40] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
57. r2_hidden(sK4(X0,X1),X0) | r1_xboole_0(X0,X1) [cnf transformation 42] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
58. r2_hidden(sK4(X0,X1),X1) | r1_xboole_0(X0,X1) [cnf transformation 42] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
59. ~r2_hidden(X2,X1) | ~r1_xboole_0(X0,X1) | ~r2_hidden(X2,X0) [cnf transformation 42] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
60. ~r2_hidden(X2,X0) | r2_hidden(X2,X1) | ~r1_tarski(X0,X1) [cnf transformation 30] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. ~r1_xboole_0(X0,X1) | r1_xboole_0(X1,X0) [cnf transformation 34] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. 1 <=> r1_xboole_0(sK0,sK1) [avatar definition]
71. r1_xboole_0(sK0,sK1) <- (1) [avatar component clause 69] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. 1 [avatar split clause 50,69]
74. 2 <=> r1_tarski(sK2,sK1) [avatar definition]
76. r1_tarski(sK2,sK1) <- (2) [avatar component clause 74] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
77. 2 [avatar split clause 49,74]
79. 3 <=> r1_tarski(sK2,sK0) [avatar definition]
81. r1_tarski(sK2,sK0) <- (3) [avatar component clause 79] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
82. 3 [avatar split clause 48,79]
84. 4 <=> v1_xboole_0(sK2) [avatar definition]
86. ~v1_xboole_0(sK2) <- (~4) [avatar component clause 84] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. ~4 [avatar split clause 47,84]
130. ~r2_hidden(sK3(X1),X0) | ~r1_xboole_0(X0,X1) | v1_xboole_0(X1) [resolution 59,56] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
132. ~r2_hidden(sK4(X7,X6),X5) | ~r1_xboole_0(X5,X6) | r1_xboole_0(X7,X6) [resolution 59,58] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
134. r2_hidden(sK3(X0),X1) | ~r1_tarski(X0,X1) | v1_xboole_0(X0) [resolution 60,56] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
135. r2_hidden(sK4(X2,X3),X4) | ~r1_tarski(X2,X4) | r1_xboole_0(X2,X3) [resolution 60,57] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
139. ~r1_tarski(X0,X1) | v1_xboole_0(X0) | ~r1_xboole_0(X1,X0) | v1_xboole_0(X0) [resolution 134,130] {a:2,w:10,wCS:4000,sine:1,thAx:0,allAx:7,thDist:-7}
147. ~r1_xboole_0(X1,X0) | v1_xboole_0(X0) | ~r1_tarski(X0,X1) [duplicate literal removal 139] {a:2,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:7,thDist:-7}
160. ~r1_tarski(X0,X1) | r1_xboole_0(X0,X2) | ~r1_xboole_0(X1,X2) | r1_xboole_0(X0,X2) [resolution 135,132] {a:2,w:12,wCS:4800,sine:1,thAx:0,allAx:7,thDist:-7}
170. ~r1_xboole_0(X1,X2) | r1_xboole_0(X0,X2) | ~r1_tarski(X0,X1) [duplicate literal removal 160] {a:2,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:7,thDist:-7}
246. ~r1_tarski(X0,sK0) | r1_xboole_0(X0,sK1) <- (1) [resolution 170,71] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
258. r1_xboole_0(sK2,sK1) <- (1, 3) [resolution 246,81] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
261. 19 <=> r1_xboole_0(sK2,sK1) [avatar definition]
263. r1_xboole_0(sK2,sK1) <- (19) [avatar component clause 261] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
264. 19 | ~1 | ~3 [avatar split clause 258,79,69,261]
270. r1_xboole_0(sK1,sK2) <- (19) [resolution 263,65] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
283. 22 <=> r1_xboole_0(sK1,sK2) [avatar definition]
285. r1_xboole_0(sK1,sK2) <- (22) [avatar component clause 283] {a:5,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:12,thDist:-12}
286. 22 | ~19 [avatar split clause 270,261,283]
316. v1_xboole_0(sK2) | ~r1_tarski(sK2,sK1) <- (22) [resolution 285,147] {a:6,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
321. ~r1_tarski(sK2,sK1) <- (~4, 22) [subsumption resolution 316,86] {a:6,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
322. $false <- (2, ~4, 22) [subsumption resolution 321,76] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
323. ~2 | 4 | ~22 [avatar contradiction clause 322]
324. $false [avatar sat refutation 72,77,82,87,264,286,323] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:262,thDist:-262}
% SZS output end Proof for xboole_1__t68_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% >>> Preprocessing
% Introduced skolems: 7
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 185
% Active clauses: 56
% Passive clauses: 106
% Final active clauses: 55
% Final passive clauses: 47
% Main loop iterations started: 59
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 7
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 27
% 
% >>> Generating Inferences
% Binary resolution: 125
% Unit resulting resolution: 1
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 18
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 28
% SAT solver unit clauses: 9
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward demodulation index maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

