i: [1,65,0,2,1,1,0,0] t14_aff_2
a: [3,90,0,2,1,65]
s: 90
i: [1,64,0,2,1,1,0,0] t14_aff_2
a: [3,95,0,2,1,64]
s: 95
i: [1,63,0,2,1,1,0,0] t14_aff_2
a: [3,100,0,2,1,63]
s: 100
i: [1,62,0,2,1,1,0,0] t14_aff_2
a: [3,105,0,2,1,62]
s: 105
i: [1,61,0,2,1,1,0,0] t14_aff_2
a: [3,110,0,2,1,61]
s: 110
i: [1,60,0,2,1,1,0,0] t14_aff_2
a: [3,115,0,2,1,60]
s: 115
i: [1,85,0,2,1,0,0,1] existence_l1_analoaf
a: [3,120,0,2,1,85]
s: 120
i: [1,86,0,2,1,0,0,2] existence_l1_struct_0
a: [3,125,0,2,1,86]
s: 125
i: [1,66,0,4,2,0,0,1] dt_l1_analoaf
s: 66
i: [1,67,0,6,3,0,0,255] cc7_struct_0
s: 67
i: [1,68,0,6,3,0,0,255] cc5_struct_0
s: 68
d: [2,127,1,2,1,1,40,66,100]
a: [3,132,1,2,1,127]
s: 132
i: [1,70,0,6,3,0,0,255] cc17_struct_0
s: 70
d: [2,128,1,2,1,1,40,66,120]
a: [3,137,1,2,1,128]
s: 137
i: [1,71,0,6,3,0,0,255] cc2_struct_0
s: 71
i: [1,72,0,6,3,0,0,255] cc1_struct_0
s: 72
d: [2,139,1,4,2,1,40,67,125]
a: [3,143,1,2,1,139]
s: 143
i: [1,75,0,6,3,0,0,255] cc4_struct_0
s: 75
d: [2,150,2,4,2,1,40,137,67]
a: [3,154,2,2,1,150]
s: 154
i: [1,76,0,6,3,0,0,255] cc6_struct_0
s: 76
i: [1,73,0,7,3,0,0,255] cc8_struct_0
s: 73
i: [1,77,0,7,3,0,0,255] cc9_struct_0
s: 77
i: [1,78,0,7,3,0,0,255] cc11_struct_0
s: 78
i: [1,79,0,7,3,0,0,255] cc11_struct_0
s: 79
i: [1,80,0,9,4,0,0,255] cc10_struct_0
s: 80
i: [1,81,0,12,6,0,0,2] t13_aff_2
s: 81
i: [1,82,0,12,6,0,0,1] t3_aff_2
s: 82
i: [1,83,0,12,6,0,0,1] t7_aff_2
s: 83
d: [2,166,1,10,5,1,40,82,95]
d: [2,167,1,8,4,2,21,166,100]
d: [2,168,1,6,3,3,21,167,105]
d: [2,169,1,4,2,4,21,168,110]
d: [2,170,1,2,1,5,21,169,115]
a: [3,174,1,2,1,170]
s: 174
i: [1,84,0,12,6,0,0,1] t7_aff_2
s: 84
d: [2,176,2,10,5,1,40,174,81]
d: [2,177,2,8,4,2,21,176,100]
d: [2,178,2,6,3,3,21,177,105]
d: [2,179,2,4,2,4,21,178,110]
d: [2,180,2,2,1,5,21,179,115]
a: [3,184,2,2,1,180]
s: 184
d: [2,186,3,0,0,6,52,115,110,105,100,90,184,84]
e: 186
d: [2,187,3,10,5,1,40,184,84]
d: [2,189,3,8,4,2,21,187,90]
d: [2,190,3,6,3,3,21,189,100]
d: [2,191,3,4,2,4,21,190,105]
d: [2,192,3,2,1,5,21,191,110]
d: [2,193,3,0,0,6,21,192,115]
e: 193
f: 193,180,170,60,61,62,63,64,65
% Refutation found. Thanks to Tanya!
% SZS status Theorem for aff_2__t14_aff_2
% SZS output start Proof for aff_2__t14_aff_2
1. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v7_aff_2(X0) => v11_aff_2(X0))) [input t14_aff_2]
2. ~! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v7_aff_2(X0) => v11_aff_2(X0))) [negated conjecture 1]
19. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v8_aff_2(X0) => v12_aff_2(X0))) [input t13_aff_2]
20. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v7_aff_2(X0) => v8_aff_2(X0))) [input t3_aff_2]
21. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v11_aff_2(X0) <=> v12_aff_2(X0))) [input t7_aff_2]
22. ? [X0] : ((~v11_aff_2(X0) & v7_aff_2(X0)) & (l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
23. ? [X0] : (~v11_aff_2(X0) & v7_aff_2(X0) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) [flattening 22]
47. ! [X0] : ((v12_aff_2(X0) | ~v8_aff_2(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 19]
48. ! [X0] : (v12_aff_2(X0) | ~v8_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 47]
49. ! [X0] : ((v8_aff_2(X0) | ~v7_aff_2(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 20]
50. ! [X0] : (v8_aff_2(X0) | ~v7_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 49]
51. ! [X0] : ((v11_aff_2(X0) <=> v12_aff_2(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 21]
52. ! [X0] : ((v11_aff_2(X0) <=> v12_aff_2(X0)) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 51]
53. ? [X0] : (~v11_aff_2(X0) & v7_aff_2(X0) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (~v11_aff_2(sK0) & v7_aff_2(sK0) & l1_analoaf(sK0) & v2_diraf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
54. ~v11_aff_2(sK0) & v7_aff_2(sK0) & l1_analoaf(sK0) & v2_diraf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0) [skolemisation 23,53]
55. ! [X0] : (((v11_aff_2(X0) | ~v12_aff_2(X0)) & (v12_aff_2(X0) | ~v11_aff_2(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [nnf transformation 52]
60. ~v7_struct_0(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
61. v1_diraf(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
62. v2_diraf(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
63. l1_analoaf(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
64. v7_aff_2(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. ~v11_aff_2(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
81. ~v8_aff_2(X0) | v12_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 48] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
82. ~v7_aff_2(X0) | v8_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 50] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
84. ~v12_aff_2(X0) | v11_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 55] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
88. 1 <=> v11_aff_2(sK0) [avatar definition]
90. ~v11_aff_2(sK0) <- (~1) [avatar component clause 88] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
91. ~1 [avatar split clause 65,88]
93. 2 <=> v7_aff_2(sK0) [avatar definition]
95. v7_aff_2(sK0) <- (2) [avatar component clause 93] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
96. 2 [avatar split clause 64,93]
98. 3 <=> l1_analoaf(sK0) [avatar definition]
100. l1_analoaf(sK0) <- (3) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
101. 3 [avatar split clause 63,98]
103. 4 <=> v2_diraf(sK0) [avatar definition]
105. v2_diraf(sK0) <- (4) [avatar component clause 103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
106. 4 [avatar split clause 62,103]
108. 5 <=> v1_diraf(sK0) [avatar definition]
110. v1_diraf(sK0) <- (5) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
111. 5 [avatar split clause 61,108]
113. 6 <=> v7_struct_0(sK0) [avatar definition]
115. ~v7_struct_0(sK0) <- (~6) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. ~6 [avatar split clause 60,113]
166. v8_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2) [resolution 82,95] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
167. v8_aff_2(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3) [subsumption resolution 166,100] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
168. v8_aff_2(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3, 4) [subsumption resolution 167,105] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
169. v8_aff_2(sK0) | v7_struct_0(sK0) <- (2, 3, 4, 5) [subsumption resolution 168,110] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
170. v8_aff_2(sK0) <- (2, 3, 4, 5, ~6) [subsumption resolution 169,115] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
172. 15 <=> v8_aff_2(sK0) [avatar definition]
174. v8_aff_2(sK0) <- (15) [avatar component clause 172] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
175. 15 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 170,113,108,103,98,93,172]
176. v12_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (15) [resolution 174,81] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
177. v12_aff_2(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 15) [subsumption resolution 176,100] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
178. v12_aff_2(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 4, 15) [subsumption resolution 177,105] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
179. v12_aff_2(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 15) [subsumption resolution 178,110] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
180. v12_aff_2(sK0) <- (3, 4, 5, ~6, 15) [subsumption resolution 179,115] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
182. 16 <=> v12_aff_2(sK0) [avatar definition]
184. v12_aff_2(sK0) <- (16) [avatar component clause 182] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
185. 16 | ~3 | ~4 | ~5 | 6 | ~15 [avatar split clause 180,172,113,108,103,98,182]
187. v11_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (16) [resolution 184,84] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
189. ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 16) [subsumption resolution 187,90] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
190. ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 3, 16) [subsumption resolution 189,100] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
191. ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 3, 4, 16) [subsumption resolution 190,105] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
192. v7_struct_0(sK0) <- (~1, 3, 4, 5, 16) [subsumption resolution 191,110] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
193. $false <- (~1, 3, 4, 5, ~6, 16) [subsumption resolution 192,115] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
194. 1 | ~3 | ~4 | ~5 | 6 | ~16 [avatar contradiction clause 193]
195. $false [avatar sat refutation 91,96,101,106,111,116,175,185,194] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:86,thDist:-86}
% SZS output end Proof for aff_2__t14_aff_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 69
% Active clauses: 27
% Passive clauses: 31
% Final active clauses: 27
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 8
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

