i: [1,65595,0,32,16,1,0,0] t3_gate_3
a: [3,65605,0,2,1,65595]
s: 65605
a: [3,65608,0,2,1,65595]
s: 65608
a: [3,65613,0,2,1,65595]
s: 65613
a: [3,65617,0,2,1,65595]
s: 65617
a: [3,65620,0,2,1,65595]
s: 65620
a: [3,65625,0,2,1,65595]
s: 65625
a: [3,65629,0,2,1,65595]
s: 65629
a: [3,65633,0,2,1,65595]
s: 65633
a: [3,65637,0,2,1,65595]
s: 65637
a: [3,65640,0,2,1,65595]
s: 65640
a: [3,65645,0,2,1,65595]
s: 65645
a: [3,65649,0,2,1,65595]
s: 65649
a: [3,65653,0,2,1,65595]
s: 65653
a: [3,65656,0,2,1,65595]
s: 65656
a: [3,65661,0,2,1,65595]
s: 65661
a: [3,65664,0,2,1,65595]
s: 65664
i: [1,59,0,4,2,1,0,0] t3_gate_3
a: [3,65925,0,2,1,59]
s: 65925
a: [3,65928,0,2,1,59]
s: 65928
i: [1,57,0,4,2,1,0,0] t3_gate_3
a: [3,65935,0,2,1,57]
s: 65935
a: [3,65938,0,2,1,57]
s: 65938
i: [1,55,0,5,2,1,0,0] t3_gate_3
a: [3,65949,0,2,1,55]
s: 65949
a: [3,65944,0,3,1,55]
s: 65944
i: [1,53,0,7,2,1,0,0] t3_gate_3
a: [3,65955,0,5,1,53]
s: 65955
i: [1,37,0,7,2,1,0,0] t3_gate_3
a: [3,66003,0,5,1,37]
s: 66003
i: [1,51,0,8,2,1,0,0] t3_gate_3
a: [3,65961,0,6,1,51]
s: 65961
i: [1,49,0,8,2,1,0,0] t3_gate_3
a: [3,65966,0,6,1,49]
s: 65966
i: [1,45,0,8,2,1,0,0] t3_gate_3
a: [3,65979,0,6,1,45]
s: 65979
i: [1,35,0,8,2,1,0,0] t3_gate_3
a: [3,66009,0,6,1,35]
s: 66009
i: [1,33,0,8,2,1,0,0] t3_gate_3
a: [3,66015,0,6,1,33]
s: 66015
i: [1,29,0,8,2,1,0,0] t3_gate_3
a: [3,66026,0,6,1,29]
s: 66026
i: [1,47,0,9,2,1,0,0] t3_gate_3
a: [3,65972,0,7,1,47]
s: 65972
i: [1,43,0,9,2,1,0,0] t3_gate_3
a: [3,65984,0,7,1,43]
s: 65984
i: [1,41,0,9,2,1,0,0] t3_gate_3
a: [3,65990,0,7,1,41]
s: 65990
i: [1,31,0,9,2,1,0,0] t3_gate_3
a: [3,66021,0,7,1,31]
s: 66021
i: [1,27,0,9,2,1,0,0] t3_gate_3
a: [3,66033,0,7,1,27]
s: 66033
i: [1,25,0,9,2,1,0,0] t3_gate_3
a: [3,66039,0,7,1,25]
s: 66039
i: [1,39,0,10,2,1,0,0] t3_gate_3
a: [3,65997,0,8,1,39]
s: 65997
i: [1,23,0,10,2,1,0,0] t3_gate_3
a: [3,66045,0,8,1,23]
s: 66045
i: [1,65596,0,5,2,0,0,1] fc2_gate_1
s: 65596
r: 65944
r: 65949
a: [3,65948,0,2,1,55]
s: 65948
a: [3,65945,0,3,1,55]
s: 65945
i: [1,65597,0,5,2,0,0,1] fc1_gate_1
s: 65597
i: [1,65598,0,7,2,0,0,1] fc20_gate_1
s: 65598
d: [2,66057,1,0,0,2,52,65948,65966,65598]
e: 66057
d: [2,66058,1,0,0,2,52,65935,66026,65598]
e: 66058
d: [2,66059,1,0,0,2,52,65948,65990,65598]
e: 66059
d: [2,66060,1,2,1,1,40,65598,65966]
d: [2,66068,1,0,0,2,21,66060,65948]
e: 66068
d: [2,66062,1,2,1,1,40,65598,66026]
d: [2,66070,1,0,0,2,21,66062,65935]
e: 66070
d: [2,66064,1,2,1,1,40,65598,65990]
d: [2,66072,1,0,0,2,21,66064,65948]
e: 66072
r: 65608
r: 65620
r: 65664
r: 65966
r: 65990
r: 66026
a: [3,65609,0,2,1,65595]
s: 65609
a: [3,65621,0,2,1,65595]
s: 65621
a: [3,65665,0,2,1,65595]
s: 65665
d: [2,66048,1,2,1,1,40,65596,65944]
a: [3,66051,1,2,1,66048]
s: 66051
a: [3,65967,0,6,1,49]
s: 65967
d: [2,66063,1,3,1,1,40,65598,65984]
a: [3,66077,1,3,1,66063]
s: 66077
a: [3,66027,0,6,1,29]
s: 66027
a: [3,65991,0,7,1,41]
s: 65991
i: [1,65599,0,7,2,0,0,1] fc18_gate_1
s: 65599
d: [2,66081,1,0,0,2,52,65928,65972,65599]
e: 66081
d: [2,66082,1,2,1,1,40,65599,65972]
d: [2,66085,1,0,0,2,21,66082,65928]
e: 66085
r: 65925
r: 65928
a: [3,65924,0,2,1,59]
s: 65924
a: [3,65929,0,2,1,59]
s: 65929
i: [1,65600,0,7,2,0,0,1] fc19_gate_1
s: 65600
d: [2,66093,1,0,0,2,52,65938,65984,65600]
e: 66093
d: [2,66095,1,2,1,1,40,65600,65984]
d: [2,66097,1,0,0,2,21,66095,65938]
e: 66097
r: 65935
r: 65938
a: [3,65934,0,2,1,57]
s: 65934
a: [3,65939,0,2,1,57]
s: 65939
i: [1,65601,0,11,4,0,0,1] fc17_gate_1
s: 65601
d: [2,66105,2,0,0,4,52,65929,65939,66077,65961,65601]
e: 66105
d: [2,66106,2,0,0,4,52,66051,65924,65934,66003,65601]
e: 66106
d: [2,66117,1,6,3,1,40,65601,66003]
d: [2,66127,1,4,2,2,21,66117,65934]
d: [2,66128,1,2,1,3,21,66127,65924]
d: [2,66129,1,0,0,4,21,66128,66051]
e: 66129
d: [2,66122,1,7,3,1,40,65601,65961]
d: [2,66136,1,4,2,2,21,66122,66077]
d: [2,66137,1,2,1,3,21,66136,65939]
d: [2,66138,1,0,0,4,21,66137,65929]
e: 66138
r: 65633
r: 65945
r: 65948
r: 65961
r: 66051
r: 65605
r: 65609
r: 65613
r: 65617
r: 65640
r: 65656
r: 65924
r: 65929
r: 65972
r: 65979
r: 65984
r: 65991
r: 66039
r: 66045
r: 66077
a: [3,65604,0,2,1,65595]
s: 65604
s: 65608
a: [3,65612,0,2,1,65595]
s: 65612
a: [3,65616,0,2,1,65595]
s: 65616
a: [3,65632,0,2,1,65595]
s: 65632
a: [3,65641,0,2,1,65595]
s: 65641
a: [3,65657,0,2,1,65595]
s: 65657
s: 65925
s: 65928
s: 65949
a: [3,66052,1,2,1,66048]
s: 66052
s: 65944
a: [3,65960,0,6,1,51]
s: 65960
d: [2,66178,1,0,0,2,52,65928,65960,65599]
e: 66178
d: [2,66180,1,2,1,1,40,65960,65599]
d: [2,66183,1,0,0,2,21,66180,65928]
e: 66183
r: 65632
r: 65653
r: 65960
r: 66015
s: 65633
a: [3,65652,0,2,1,65595]
s: 65652
a: [3,66076,1,3,1,66063]
s: 66076
s: 65961
a: [3,65978,0,6,1,45]
s: 65978
d: [2,66083,1,3,1,1,40,65599,65984]
a: [3,66091,1,3,1,66083]
s: 66091
a: [3,66014,0,6,1,33]
s: 66014
d: [2,66192,2,0,0,2,52,66052,66014,65599]
e: 66192
d: [2,66194,1,2,1,1,40,66014,65599]
d: [2,66197,1,0,0,2,21,66194,66052]
e: 66197
r: 65645
r: 65652
r: 66014
r: 66021
a: [3,65644,0,2,1,65595]
s: 65644
s: 65653
s: 66015
d: [2,66094,1,3,1,1,40,65600,65972]
a: [3,66103,1,3,1,66094]
s: 66103
d: [2,66200,2,0,0,2,52,65939,66103,65596]
e: 66200
d: [2,66201,2,2,1,1,40,66103,65596]
d: [2,66203,2,0,0,2,21,66201,65939]
e: 66203
r: 65934
r: 65939
r: 65616
r: 65978
s: 65617
s: 65935
s: 65938
s: 65979
a: [3,65973,0,7,1,47]
s: 65973
a: [3,65985,0,7,1,43]
s: 65985
s: 65990
a: [3,66020,0,7,1,31]
s: 66020
d: [2,66213,2,0,0,2,52,66052,66020,65599]
e: 66213
d: [2,66215,1,2,1,1,40,66020,65599]
d: [2,66218,1,0,0,2,21,66215,66052]
e: 66218
r: 65644
r: 66020
s: 65645
s: 66021
a: [3,66038,0,7,1,25]
s: 66038
d: [2,66221,1,0,0,2,52,65935,66038,65598]
e: 66221
d: [2,66224,1,2,1,1,40,66038,65598]
d: [2,66226,1,0,0,2,21,66224,65935]
e: 66226
r: 65612
r: 65629
r: 66003
r: 66038
s: 65613
a: [3,65628,0,2,1,65595]
s: 65628
a: [3,66002,0,5,1,37]
s: 66002
d: [2,66228,1,0,0,2,52,65925,66002,65600]
e: 66228
d: [2,66229,2,0,0,2,52,66052,66002,65599]
e: 66229
d: [2,66230,1,0,0,2,52,65935,66002,65598]
e: 66230
d: [2,66231,1,2,1,1,40,66002,65600]
d: [2,66237,1,0,0,2,21,66231,65925]
e: 66237
d: [2,66232,1,2,1,1,40,66002,65599]
d: [2,66239,1,0,0,2,21,66232,66052]
e: 66239
d: [2,66233,1,2,1,1,40,66002,65598]
d: [2,66241,1,0,0,2,21,66233,65935]
e: 66241
r: 65628
r: 65637
r: 66002
r: 66009
s: 65629
a: [3,65636,0,2,1,65595]
s: 65636
s: 66003
a: [3,66008,0,6,1,35]
s: 66008
d: [2,66244,1,0,0,2,52,65925,66008,65600]
e: 66244
d: [2,66245,2,0,0,2,52,66052,66008,65599]
e: 66245
d: [2,66246,1,2,1,1,40,66008,65600]
d: [2,66251,1,0,0,2,21,66246,65925]
e: 66251
d: [2,66247,1,2,1,1,40,66008,65599]
d: [2,66253,1,0,0,2,21,66247,66052]
e: 66253
r: 65636
r: 65665
r: 65967
r: 66008
s: 65637
s: 65664
s: 65966
d: [2,66255,1,0,0,2,52,65928,65966,65599]
e: 66255
d: [2,66257,1,2,1,1,40,65966,65599]
d: [2,66260,1,0,0,2,21,66257,65928]
e: 66260
r: 65925
r: 65928
r: 65604
r: 65664
r: 65966
s: 65605
s: 65665
s: 65925
s: 65928
s: 65967
d: [2,66115,1,9,3,1,40,65601,66045]
a: [3,66147,1,3,1,66115]
s: 66147
d: [2,66265,2,0,0,2,52,65935,66147,65597]
e: 66265
d: [2,66266,2,2,1,1,40,66147,65597]
d: [2,66268,2,0,0,2,21,66266,65935]
e: 66268
r: 66147
s: 66009
s: 66039
a: [3,66143,1,3,1,66115]
s: 66143
d: [2,66272,2,0,0,2,52,65925,66143,65597]
e: 66272
d: [2,66273,2,2,1,1,40,66143,65597]
d: [2,66275,2,0,0,2,21,66273,65925]
e: 66275
r: 65605
r: 65625
r: 65955
r: 66143
s: 65604
a: [3,65624,0,2,1,65595]
s: 65624
a: [3,65954,0,5,1,53]
s: 65954
d: [2,66277,1,0,0,2,52,65938,65954,65600]
e: 66277
d: [2,66278,1,0,0,2,52,65928,65954,65599]
e: 66278
d: [2,66279,1,2,1,1,40,65954,65600]
d: [2,66284,1,0,0,2,21,66279,65938]
e: 66284
d: [2,66280,1,2,1,1,40,65954,65599]
d: [2,66286,1,0,0,2,21,66280,65928]
e: 66286
r: 65624
r: 65661
r: 65954
r: 66033
s: 65625
a: [3,65660,0,2,1,65595]
s: 65660
s: 65955
a: [3,66032,0,7,1,27]
s: 66032
d: [2,66289,1,0,0,2,52,65925,66032,65600]
e: 66289
d: [2,66290,1,2,1,1,40,66032,65600]
d: [2,66294,1,0,0,2,21,66290,65925]
e: 66294
r: 65608
r: 65645
r: 65660
r: 65944
r: 65949
r: 65990
r: 66021
r: 66032
r: 66052
r: 65649
r: 65997
r: 66076
s: 65609
s: 65644
a: [3,65648,0,2,1,65595]
s: 65648
s: 65661
s: 65948
s: 65945
s: 65991
s: 66051
s: 66020
s: 66077
s: 66033
a: [3,65996,0,8,1,39]
s: 65996
a: [3,66142,1,3,1,66115]
s: 66142
a: [3,66044,0,8,1,23]
s: 66044
d: [2,66311,1,0,0,2,52,65945,66044,65599]
e: 66311
d: [2,66313,1,3,1,1,40,66044,65599]
d: [2,66316,1,0,0,2,21,66313,65945]
e: 66316
r: 65604
r: 65653
r: 66015
r: 66044
s: 65605
s: 65652
s: 66014
d: [2,66318,1,0,0,2,52,65935,66014,65598]
e: 66318
d: [2,66321,1,2,1,1,40,66014,65598]
d: [2,66323,1,0,0,2,21,66321,65935]
e: 66323
r: 65605
r: 65609
r: 65644
r: 65652
r: 65945
r: 65948
r: 65991
r: 66014
r: 66020
r: 66051
r: 66077
r: 65648
r: 65661
r: 65665
r: 65925
r: 65928
r: 65967
r: 65996
r: 66033
r: 66091
r: 66142
s: 65605
s: 65608
s: 65645
s: 65649
s: 65653
s: 65660
s: 65664
s: 65924
s: 65929
s: 65949
s: 66052
s: 65944
s: 66076
s: 65966
s: 66015
a: [3,66090,1,3,1,66083]
s: 66090
s: 65990
d: [2,66336,2,0,0,2,52,66090,65990,65599]
e: 66336
d: [2,66338,1,3,1,1,40,65990,65599]
d: [2,66341,1,0,0,2,21,66338,66090]
e: 66341
r: 65605
r: 65608
r: 65990
s: 65604
s: 65609
s: 65991
s: 66143
s: 66021
a: [3,66146,1,3,1,66115]
s: 66146
s: 66032
s: 65997
s: 66044
d: [2,66351,2,0,0,2,52,66143,66044,65600]
e: 66351
d: [2,66352,1,3,1,1,40,66044,65600]
d: [2,66356,1,0,0,2,21,66352,66143]
e: 66356
r: 65604
r: 65633
r: 65961
r: 66044
s: 65605
s: 65632
s: 65960
d: [2,66358,1,0,0,2,52,65938,65960,65600]
e: 66358
d: [2,66359,2,0,0,2,52,66076,65960,65598]
e: 66359
d: [2,66360,1,2,1,1,40,65960,65600]
d: [2,66365,1,0,0,2,21,66360,65938]
e: 66365
d: [2,66362,1,3,1,1,40,65960,65598]
d: [2,66367,1,0,0,2,21,66362,66076]
e: 66367
r: 65605
r: 65632
r: 65960
r: 66143
s: 65604
s: 65633
s: 65961
s: 66044
s: 66142
d: [2,66374,2,0,0,2,52,65924,66142,65596]
e: 66374
d: [2,66375,2,2,1,1,40,66142,65596]
d: [2,66377,2,0,0,2,21,66375,65924]
e: 66377
r: 65660
r: 65664
r: 65924
r: 65929
r: 65966
r: 66032
r: 65604
r: 65637
r: 65641
r: 65649
r: 65944
r: 65949
r: 65973
r: 65997
r: 66009
r: 66044
r: 66052
r: 66076
r: 66142
s: 65605
s: 65636
s: 65640
s: 65648
s: 65661
s: 65665
s: 65924
s: 65929
s: 65948
s: 65945
s: 65967
s: 66051
s: 66008
s: 66077
s: 65972
s: 66033
s: 66143
s: 65996
d: [2,66395,2,0,0,2,52,66090,65996,65599]
e: 66395
d: [2,66397,1,3,1,1,40,65996,65599]
d: [2,66400,1,0,0,2,21,66397,66090]
e: 66400
r: 65648
r: 65661
r: 65996
r: 66033
s: 65649
s: 65660
s: 66032
d: [2,66402,1,0,0,2,52,65945,66032,65599]
e: 66402
d: [2,66404,1,3,1,1,40,66032,65599]
d: [2,66407,1,0,0,2,21,66404,65945]
e: 66407
r: 65640
r: 65660
r: 65972
r: 66032
r: 66103
s: 65641
s: 65661
a: [3,66102,1,3,1,66094]
s: 66102
d: [2,66409,2,0,0,2,52,65938,66102,65597]
e: 66409
d: [2,66410,2,2,1,1,40,66102,65597]
d: [2,66412,2,0,0,2,21,66410,65938]
e: 66412
r: 65629
r: 65633
r: 65935
r: 65938
r: 65961
r: 66003
r: 66146
r: 65625
r: 65636
r: 65955
r: 66008
s: 65624
s: 65628
s: 65632
s: 65637
s: 65934
s: 65939
s: 65954
d: [2,66417,1,0,0,2,52,65948,65954,65598]
e: 66417
d: [2,66420,1,2,1,1,40,65954,65598]
d: [2,66422,1,0,0,2,21,66420,65948]
e: 66422
r: 65613
r: 65624
r: 65954
r: 66039
s: 65612
s: 65625
s: 66147
s: 65955
s: 66002
s: 65960
s: 66009
s: 65973
s: 66033
s: 66038
d: [2,66435,2,0,0,2,52,66143,66038,65600]
e: 66435
d: [2,66436,1,0,0,2,52,65945,66038,65599]
e: 66436
d: [2,66437,1,3,1,1,40,66038,65600]
d: [2,66442,1,0,0,2,21,66437,66143]
e: 66442
d: [2,66438,1,3,1,1,40,66038,65599]
d: [2,66444,1,0,0,2,21,66438,65945]
e: 66444
r: 65612
r: 65621
r: 66027
r: 66038
s: 65613
s: 65620
s: 66026
d: [2,66446,1,0,0,2,52,65945,66026,65599]
e: 66446
d: [2,66448,1,3,1,1,40,66026,65599]
d: [2,66451,1,0,0,2,21,66448,65945]
e: 66451
r: 65620
r: 65653
r: 66015
r: 66026
s: 65621
s: 65652
s: 66014
d: [2,66453,2,0,0,2,52,66143,66014,65600]
e: 66453
d: [2,66454,1,3,1,1,40,66014,65600]
d: [2,66458,1,0,0,2,21,66454,66143]
e: 66458
r: 65628
r: 65632
r: 65924
r: 65929
r: 65960
r: 66002
r: 66090
r: 66143
r: 65621
r: 65625
r: 65652
r: 65665
r: 65945
r: 65948
r: 65955
r: 65967
r: 66014
r: 66051
r: 66077
s: 65620
s: 65624
s: 65629
s: 65633
s: 65653
s: 65664
s: 65924
s: 65929
s: 65949
s: 66052
s: 65944
s: 66076
s: 65954
s: 66003
s: 66090
s: 65961
s: 66143
s: 65966
d: [2,66477,2,0,0,2,52,66102,65966,65600]
e: 66477
d: [2,66478,1,3,1,1,40,65966,65600]
d: [2,66482,1,0,0,2,21,66478,66102]
e: 66482
r: 65613
r: 65617
r: 65624
r: 65664
r: 65924
r: 65929
r: 65954
r: 65966
r: 65979
r: 66090
r: 66143
s: 65612
s: 65616
s: 65625
s: 65665
s: 65925
s: 65928
s: 66091
s: 65955
s: 65967
s: 65978
s: 66015
s: 66026
d: [2,66491,1,0,0,2,52,65925,66026,65600]
e: 66491
d: [2,66492,1,2,1,1,40,66026,65600]
d: [2,66496,1,0,0,2,21,66492,65925]
e: 66496
r: 65605
r: 65620
r: 66026
s: 65604
s: 65621
s: 66027
s: 66038
s: 66142
s: 65997
s: 66044
d: [2,66505,2,0,0,2,52,66147,66044,65598]
e: 66505
d: [2,66508,1,3,1,1,40,66044,65598]
d: [2,66510,1,0,0,2,21,66508,66147]
e: 66510
r: 65604
r: 65609
r: 65991
r: 66044
s: 65605
s: 65608
s: 65990
d: [2,66512,2,0,0,2,52,66102,65990,65600]
e: 66512
d: [2,66513,1,3,1,1,40,65990,65600]
d: [2,66517,1,0,0,2,21,66513,66102]
e: 66517
r: 65608
r: 65612
r: 65616
r: 65621
r: 65625
r: 65661
r: 65925
r: 65928
r: 65955
r: 65978
r: 65990
r: 66027
r: 66033
r: 66038
r: 66091
r: 66142
s: 65609
s: 65613
s: 65617
s: 65620
s: 65624
s: 65660
s: 65924
s: 65929
s: 65954
s: 66090
s: 65979
s: 66143
s: 66026
s: 65991
s: 66032
d: [2,66529,2,0,0,2,52,66147,66032,65598]
e: 66529
d: [2,66532,1,3,1,1,40,66032,65598]
d: [2,66534,1,0,0,2,21,66532,66147]
e: 66534
r: 65617
r: 65620
r: 65624
r: 65653
r: 65657
r: 65660
r: 65924
r: 65929
r: 65944
r: 65949
r: 65954
r: 65979
r: 65985
r: 66015
r: 66026
r: 66032
r: 66052
r: 66076
r: 66090
r: 66143
s: 65616
s: 65621
s: 65625
s: 65652
s: 65656
s: 65661
s: 65925
s: 65928
s: 65948
s: 66051
s: 65945
s: 66077
s: 65955
s: 65978
d: [2,66543,1,0,0,2,52,65948,65978,65598]
e: 66543
d: [2,66546,1,2,1,1,40,65978,65598]
d: [2,66548,1,0,0,2,21,66546,65948]
e: 66548
r: 65616
r: 65649
r: 65978
r: 65997
s: 65617
s: 65648
s: 66091
s: 65979
s: 66014
s: 66142
s: 66027
s: 65984
s: 66033
s: 66039
s: 65996
d: [2,66563,2,0,0,2,52,66102,65996,65600]
e: 66563
d: [2,66564,1,3,1,1,40,65996,65600]
d: [2,66568,1,0,0,2,21,66564,66102]
e: 66568
i: [1,49212,0,32,16,1,0,0] t3_gate_3
i: [1,12348,0,32,16,1,0,0] t3_gate_3
i: [1,3132,0,32,16,1,0,0] t3_gate_3
i: [1,828,0,32,16,1,0,0] t3_gate_3
i: [1,252,0,32,16,1,0,0] t3_gate_3
i: [1,108,0,32,16,1,0,0] t3_gate_3
i: [1,72,0,32,16,1,0,0] t3_gate_3
i: [1,63,0,32,16,1,0,0] t3_gate_3
i: [1,58,0,4,2,1,0,0] t3_gate_3
i: [1,56,0,4,2,1,0,0] t3_gate_3
i: [1,54,0,5,2,1,0,0] t3_gate_3
i: [1,52,0,7,2,1,0,0] t3_gate_3
i: [1,50,0,8,2,1,0,0] t3_gate_3
i: [1,48,0,8,2,1,0,0] t3_gate_3
i: [1,46,0,9,2,1,0,0] t3_gate_3
i: [1,44,0,8,2,1,0,0] t3_gate_3
i: [1,42,0,9,2,1,0,0] t3_gate_3
i: [1,40,0,9,2,1,0,0] t3_gate_3
i: [1,38,0,10,2,1,0,0] t3_gate_3
i: [1,36,0,7,2,1,0,0] t3_gate_3
i: [1,34,0,8,2,1,0,0] t3_gate_3
i: [1,32,0,8,2,1,0,0] t3_gate_3
i: [1,30,0,9,2,1,0,0] t3_gate_3
i: [1,28,0,8,2,1,0,0] t3_gate_3
i: [1,26,0,9,2,1,0,0] t3_gate_3
i: [1,24,0,9,2,1,0,0] t3_gate_3
i: [1,22,0,10,2,1,0,0] t3_gate_3
d: [2,66061,1,3,1,1,40,65598,65972]
d: [2,66113,1,8,3,1,40,65601,66039]
d: [2,66149,1,6,2,2,21,66113,65934]
d: [2,66116,1,8,3,1,40,65601,65991]
d: [2,66080,2,2,1,1,40,66077,65596]
d: [2,66119,1,7,3,1,40,65601,66015]
d: [2,66133,1,5,2,2,21,66119,65934]
d: [2,66111,1,8,3,1,40,65601,66033]
d: [2,66112,1,7,3,1,40,65601,66027]
d: [2,66151,1,5,2,2,21,66112,65934]
d: [2,66121,1,6,3,1,40,65601,65955]
d: [2,66158,1,4,2,2,21,66121,65939]
d: [2,66123,1,7,3,1,40,65601,65967]
d: [2,66114,1,9,3,1,40,65601,65997]
d: [2,66188,1,3,1,1,40,65978,65599]
d: [2,66187,1,2,1,1,40,65978,65600]
d: [2,66191,2,2,1,1,40,66091,65596]
d: [2,66185,2,2,1,1,40,66076,65597]
d: [2,66300,1,3,1,1,40,66020,65600]
d: [2,66302,1,3,1,1,40,66020,65598]
d: [2,66309,1,3,1,1,40,65996,65598]
d: [2,66270,1,7,3,1,40,66009,65601]
d: [2,66208,1,8,3,1,40,65973,65601]
d: [2,66327,1,5,2,2,21,66208,66103]
d: [2,66370,1,3,1,1,40,66044,65600]
d: [2,66335,2,2,1,1,40,66090,65597]
d: [2,66345,1,8,3,1,40,66021,65601]
d: [2,66381,1,5,2,2,21,66345,66146]
d: [2,66383,1,2,1,1,40,65945,65597]
d: [2,66346,2,2,1,1,40,66146,65596]
d: [2,66388,1,3,1,1,40,66008,65598]
d: [2,66207,1,7,3,1,40,65979,65601]
d: [2,66464,1,5,2,2,21,66207,65939]
d: [2,66209,1,8,3,1,40,65985,65601]
d: [2,66466,1,6,2,2,21,66209,65939]
d: [2,66499,1,3,1,1,40,66038,65600]
d: [2,66473,1,6,3,1,40,66003,65601]
d: [2,66536,1,4,2,2,21,66473,65934]
d: [2,66475,1,7,3,1,40,65961,65601]
d: [2,66538,1,5,2,2,21,66475,65939]
d: [2,66552,1,3,1,1,40,66014,65600]
f: 66568,66552,66548,66538,66536,66534,66517,66510,66499,66496,66482,66466,66464,66451,66444,66422,66388,66410,66346,66407,66400,66383,66381,66335,66375,66370,66367,66365,66341,66327,66270,66309,66323,66316,66302,66300,66185,66294,66286,66284,66275,66268,66191,66257,66253,66251,66241,66239,66237,66226,66218,66187,66201,66197,66188,66183,66114,66123,66158,66151,66111,66133,66080,66116,66149,66115,66094,66095,66083,66082,66061,66063,66072,66070,66068,66048,22,23,24,25,26,27,28,29,30,31,32,33,34,35,36,37,38,39,40,41,42,43,44,45,46,47,48,49,50,51,52,53,54,55,56,57,58,59,63,72,108,252,828,3132,12348,49212
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_3__t3_gate_3
% SZS output start Proof for gate_3__t3_gate_3
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ~(~(~(v1_xboole_0(X13) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(X17)) & ~(v1_xboole_0(X17) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(X16)) & ~(v1_xboole_0(X16) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k1_gate_1(X18))) & ~(v1_xboole_0(k1_gate_1(X18)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [input t3_gate_3]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ~(~(~(v1_xboole_0(X13) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(X17)) & ~(v1_xboole_0(X17) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(X16)) & ~(v1_xboole_0(X16) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k1_gate_1(X18))) & ~(v1_xboole_0(k1_gate_1(X18)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [negated conjecture 1]
5. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc17_gate_1]
6. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc18_gate_1]
7. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X0,X2))) [input fc19_gate_1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X2,X0))) [input fc20_gate_1]
10. ! [X0] : (~v1_xboole_0(X0) => v1_xboole_0(k1_gate_1(X0))) [input fc2_gate_1]
11. ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : ~(~(~(v1_xboole_0(X13) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(X17)) & ~(v1_xboole_0(X17) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(X16)) & ~(v1_xboole_0(X16) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k1_gate_1(X18))) & ~(v1_xboole_0(k1_gate_1(X18)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [flattening 2]
12. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (((v1_xboole_0(X13) & ~v1_xboole_0(X2)) | (v1_xboole_0(X2) & ~v1_xboole_0(X13)) | (v1_xboole_0(X10) & ~v1_xboole_0(X5)) | (v1_xboole_0(X5) & ~v1_xboole_0(X10)) | (v1_xboole_0(X8) & ~v1_xboole_0(X4)) | (v1_xboole_0(X4) & ~v1_xboole_0(X8)) | (v1_xboole_0(X12) & ~v1_xboole_0(X6)) | (v1_xboole_0(X6) & ~v1_xboole_0(X12)) | (v1_xboole_0(X14) & ~v1_xboole_0(X7)) | (v1_xboole_0(X7) & ~v1_xboole_0(X14)) | (v1_xboole_0(X15) & ~v1_xboole_0(X3)) | (v1_xboole_0(X3) & ~v1_xboole_0(X15)) | (v1_xboole_0(X11) & ~v1_xboole_0(X1)) | (v1_xboole_0(X1) & ~v1_xboole_0(X11)) | (v1_xboole_0(X9) & ~v1_xboole_0(X0)) | (v1_xboole_0(X0) & ~v1_xboole_0(X9))) & (~v1_xboole_0(X21) | v1_xboole_0(X17)) & (~v1_xboole_0(X17) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(X16)) & (~v1_xboole_0(X16) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k1_gate_1(X18))) & (~v1_xboole_0(k1_gate_1(X18)) | v1_xboole_0(X19)) & (~v1_xboole_0(X15) | v1_xboole_0(k8_gate_1(X21,X20,X19))) & (~v1_xboole_0(k8_gate_1(X21,X20,X19)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) | v1_xboole_0(X10)) & (~v1_xboole_0(X9) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) | v1_xboole_0(X9)) & (~v1_xboole_0(X8) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X8)) & (~v1_xboole_0(X7) | v1_xboole_0(k8_gate_1(X18,X17,X16))) & (~v1_xboole_0(k8_gate_1(X18,X17,X16)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X0))) [ennf transformation 11]
13. ! [X0] : (v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0)) [ennf transformation 10]
14. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
15. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0)) [ennf transformation 9]
16. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0)) [ennf transformation 6]
17. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0)) [ennf transformation 7]
18. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 5]
19. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 18]
20. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (((v1_xboole_0(X13) & ~v1_xboole_0(X2)) | (v1_xboole_0(X2) & ~v1_xboole_0(X13)) | (v1_xboole_0(X10) & ~v1_xboole_0(X5)) | (v1_xboole_0(X5) & ~v1_xboole_0(X10)) | (v1_xboole_0(X8) & ~v1_xboole_0(X4)) | (v1_xboole_0(X4) & ~v1_xboole_0(X8)) | (v1_xboole_0(X12) & ~v1_xboole_0(X6)) | (v1_xboole_0(X6) & ~v1_xboole_0(X12)) | (v1_xboole_0(X14) & ~v1_xboole_0(X7)) | (v1_xboole_0(X7) & ~v1_xboole_0(X14)) | (v1_xboole_0(X15) & ~v1_xboole_0(X3)) | (v1_xboole_0(X3) & ~v1_xboole_0(X15)) | (v1_xboole_0(X11) & ~v1_xboole_0(X1)) | (v1_xboole_0(X1) & ~v1_xboole_0(X11)) | (v1_xboole_0(X9) & ~v1_xboole_0(X0)) | (v1_xboole_0(X0) & ~v1_xboole_0(X9))) & (~v1_xboole_0(X21) | v1_xboole_0(X17)) & (~v1_xboole_0(X17) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(X16)) & (~v1_xboole_0(X16) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k1_gate_1(X18))) & (~v1_xboole_0(k1_gate_1(X18)) | v1_xboole_0(X19)) & (~v1_xboole_0(X15) | v1_xboole_0(k8_gate_1(X21,X20,X19))) & (~v1_xboole_0(k8_gate_1(X21,X20,X19)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) | v1_xboole_0(X10)) & (~v1_xboole_0(X9) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) | v1_xboole_0(X9)) & (~v1_xboole_0(X8) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X8)) & (~v1_xboole_0(X7) | v1_xboole_0(k8_gate_1(X18,X17,X16))) & (~v1_xboole_0(k8_gate_1(X18,X17,X16)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X0))) => (((v1_xboole_0(sK13) & ~v1_xboole_0(sK2)) | (v1_xboole_0(sK2) & ~v1_xboole_0(sK13)) | (v1_xboole_0(sK10) & ~v1_xboole_0(sK5)) | (v1_xboole_0(sK5) & ~v1_xboole_0(sK10)) | (v1_xboole_0(sK8) & ~v1_xboole_0(sK4)) | (v1_xboole_0(sK4) & ~v1_xboole_0(sK8)) | (v1_xboole_0(sK12) & ~v1_xboole_0(sK6)) | (v1_xboole_0(sK6) & ~v1_xboole_0(sK12)) | (v1_xboole_0(sK14) & ~v1_xboole_0(sK7)) | (v1_xboole_0(sK7) & ~v1_xboole_0(sK14)) | (v1_xboole_0(sK15) & ~v1_xboole_0(sK3)) | (v1_xboole_0(sK3) & ~v1_xboole_0(sK15)) | (v1_xboole_0(sK11) & ~v1_xboole_0(sK1)) | (v1_xboole_0(sK1) & ~v1_xboole_0(sK11)) | (v1_xboole_0(sK9) & ~v1_xboole_0(sK0)) | (v1_xboole_0(sK0) & ~v1_xboole_0(sK9))) & (~v1_xboole_0(sK21) | v1_xboole_0(sK17)) & (~v1_xboole_0(sK17) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(sK16)) & (~v1_xboole_0(sK16) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK18))) & (~v1_xboole_0(k1_gate_1(sK18)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0))) [choice axiom]
21. ((v1_xboole_0(sK13) & ~v1_xboole_0(sK2)) | (v1_xboole_0(sK2) & ~v1_xboole_0(sK13)) | (v1_xboole_0(sK10) & ~v1_xboole_0(sK5)) | (v1_xboole_0(sK5) & ~v1_xboole_0(sK10)) | (v1_xboole_0(sK8) & ~v1_xboole_0(sK4)) | (v1_xboole_0(sK4) & ~v1_xboole_0(sK8)) | (v1_xboole_0(sK12) & ~v1_xboole_0(sK6)) | (v1_xboole_0(sK6) & ~v1_xboole_0(sK12)) | (v1_xboole_0(sK14) & ~v1_xboole_0(sK7)) | (v1_xboole_0(sK7) & ~v1_xboole_0(sK14)) | (v1_xboole_0(sK15) & ~v1_xboole_0(sK3)) | (v1_xboole_0(sK3) & ~v1_xboole_0(sK15)) | (v1_xboole_0(sK11) & ~v1_xboole_0(sK1)) | (v1_xboole_0(sK1) & ~v1_xboole_0(sK11)) | (v1_xboole_0(sK9) & ~v1_xboole_0(sK0)) | (v1_xboole_0(sK0) & ~v1_xboole_0(sK9))) & (~v1_xboole_0(sK21) | v1_xboole_0(sK17)) & (~v1_xboole_0(sK17) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(sK16)) & (~v1_xboole_0(sK16) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK18))) & (~v1_xboole_0(k1_gate_1(sK18)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0)) [skolemisation 12,20]
22. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0) [cnf transformation 21] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
23. ~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) [cnf transformation 21] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
24. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
25. ~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
26. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. ~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. ~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. ~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. ~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. ~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. ~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8) [cnf transformation 21] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. ~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) [cnf transformation 21] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
41. ~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. ~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
44. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. ~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. ~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. ~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. ~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. ~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. ~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. ~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. ~v1_xboole_0(k1_gate_1(sK18)) | v1_xboole_0(sK19) [cnf transformation 21] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. ~v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK18)) [cnf transformation 21] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
56. ~v1_xboole_0(sK16) | v1_xboole_0(sK20) [cnf transformation 21] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
57. ~v1_xboole_0(sK20) | v1_xboole_0(sK16) [cnf transformation 21] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. ~v1_xboole_0(sK17) | v1_xboole_0(sK21) [cnf transformation 21] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. ~v1_xboole_0(sK21) | v1_xboole_0(sK17) [cnf transformation 21] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
63. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK7) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK11) | v1_xboole_0(sK9) | v1_xboole_0(sK0) [cnf transformation 21] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
72. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK7) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK15) | v1_xboole_0(sK11) | v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 21] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
108. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK7) | ~v1_xboole_0(sK14) | v1_xboole_0(sK15) | v1_xboole_0(sK3) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 21] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
252. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK12) | v1_xboole_0(sK14) | v1_xboole_0(sK7) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 21] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
828. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK8) | v1_xboole_0(sK12) | v1_xboole_0(sK6) | ~v1_xboole_0(sK7) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 21] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
3132. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK10) | v1_xboole_0(sK8) | v1_xboole_0(sK4) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK7) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 21] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
12348. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK13) | v1_xboole_0(sK10) | v1_xboole_0(sK5) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK7) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 21] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49212. v1_xboole_0(sK13) | v1_xboole_0(sK2) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK7) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 21] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65596. v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0) [cnf transformation 13] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65597. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 14] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65598. v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0) [cnf transformation 15] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65599. v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0) [cnf transformation 16] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65600. v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0) [cnf transformation 17] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65601. ~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65603. 1 <=> v1_xboole_0(sK0) [avatar definition]
65607. 2 <=> v1_xboole_0(sK9) [avatar definition]
65611. 3 <=> v1_xboole_0(sK1) [avatar definition]
65615. 4 <=> v1_xboole_0(sK11) [avatar definition]
65619. 5 <=> v1_xboole_0(sK3) [avatar definition]
65623. 6 <=> v1_xboole_0(sK15) [avatar definition]
65627. 7 <=> v1_xboole_0(sK7) [avatar definition]
65631. 8 <=> v1_xboole_0(sK14) [avatar definition]
65635. 9 <=> v1_xboole_0(sK6) [avatar definition]
65639. 10 <=> v1_xboole_0(sK12) [avatar definition]
65643. 11 <=> v1_xboole_0(sK4) [avatar definition]
65647. 12 <=> v1_xboole_0(sK8) [avatar definition]
65651. 13 <=> v1_xboole_0(sK5) [avatar definition]
65655. 14 <=> v1_xboole_0(sK10) [avatar definition]
65659. 15 <=> v1_xboole_0(sK2) [avatar definition]
65663. 16 <=> v1_xboole_0(sK13) [avatar definition]
65793. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | 15 | 16 [avatar split clause 49212,65663,65659,65651,65655,65643,65647,65635,65639,65627,65631,65619,65623,65611,65615,65603,65607]
65857. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | 13 | 14 | ~16 | ~15 [avatar split clause 12348,65659,65663,65655,65651,65643,65647,65635,65639,65627,65631,65619,65623,65611,65615,65603,65607]
65889. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | 11 | 12 | ~14 | ~13 | ~16 | ~15 [avatar split clause 3132,65659,65663,65651,65655,65647,65643,65635,65639,65627,65631,65619,65623,65611,65615,65603,65607]
65905. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | 9 | 10 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 828,65659,65663,65651,65655,65643,65647,65639,65635,65627,65631,65619,65623,65611,65615,65603,65607]
65913. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | 7 | 8 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 252,65659,65663,65651,65655,65643,65647,65635,65639,65631,65627,65619,65623,65611,65615,65603,65607]
65917. ~2 | ~1 | ~4 | ~3 | 5 | 6 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 108,65659,65663,65651,65655,65643,65647,65635,65639,65627,65631,65623,65619,65611,65615,65603,65607]
65919. ~2 | ~1 | 3 | 4 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 72,65659,65663,65651,65655,65643,65647,65635,65639,65627,65631,65619,65623,65615,65611,65603,65607]
65920. 1 | 2 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 63,65659,65663,65651,65655,65643,65647,65635,65639,65627,65631,65619,65623,65611,65615,65607,65603]
65923. 17 <=> v1_xboole_0(sK17) [avatar definition]
65925. v1_xboole_0(sK17) <- (17) [avatar component clause 65923] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65927. 18 <=> v1_xboole_0(sK21) [avatar definition]
65928. v1_xboole_0(sK21) <- (18) [avatar component clause 65927] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65930. 17 | ~18 [avatar split clause 59,65927,65923]
65931. 18 | ~17 [avatar split clause 58,65923,65927]
65933. 19 <=> v1_xboole_0(sK16) [avatar definition]
65934. ~v1_xboole_0(sK16) <- (~19) [avatar component clause 65933] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65935. v1_xboole_0(sK16) <- (19) [avatar component clause 65933] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65937. 20 <=> v1_xboole_0(sK20) [avatar definition]
65938. v1_xboole_0(sK20) <- (20) [avatar component clause 65937] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65939. ~v1_xboole_0(sK20) <- (~20) [avatar component clause 65937] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65940. 19 | ~20 [avatar split clause 57,65937,65933]
65941. 20 | ~19 [avatar split clause 56,65933,65937]
65943. 21 <=> v1_xboole_0(k1_gate_1(sK18)) [avatar definition]
65944. ~v1_xboole_0(k1_gate_1(sK18)) <- (~21) [avatar component clause 65943] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65945. v1_xboole_0(k1_gate_1(sK18)) <- (21) [avatar component clause 65943] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65947. 22 <=> v1_xboole_0(sK19) [avatar definition]
65948. v1_xboole_0(sK19) <- (22) [avatar component clause 65947] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65950. 21 | ~22 [avatar split clause 55,65947,65943]
65951. 22 | ~21 [avatar split clause 54,65943,65947]
65953. 23 <=> v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) [avatar definition]
65954. ~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) <- (~23) [avatar component clause 65953] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65955. v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) <- (23) [avatar component clause 65953] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65956. 23 | ~6 [avatar split clause 53,65623,65953]
65957. 6 | ~23 [avatar split clause 52,65953,65623]
65959. 24 <=> v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) [avatar definition]
65960. ~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) <- (~24) [avatar component clause 65959] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65961. v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) <- (24) [avatar component clause 65959] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65962. 24 | ~8 [avatar split clause 51,65631,65959]
65963. 8 | ~24 [avatar split clause 50,65959,65631]
65965. 25 <=> v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) [avatar definition]
65966. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) <- (~25) [avatar component clause 65965] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65967. v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) <- (25) [avatar component clause 65965] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65968. 25 | ~16 [avatar split clause 49,65663,65965]
65969. 16 | ~25 [avatar split clause 48,65965,65663]
65971. 26 <=> v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) [avatar definition]
65972. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) <- (~26) [avatar component clause 65971] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65973. v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) <- (26) [avatar component clause 65971] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65974. 26 | ~10 [avatar split clause 47,65639,65971]
65975. 10 | ~26 [avatar split clause 46,65971,65639]
65977. 27 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) [avatar definition]
65978. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) <- (~27) [avatar component clause 65977] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65979. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) <- (27) [avatar component clause 65977] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65980. 27 | ~4 [avatar split clause 45,65615,65977]
65981. 4 | ~27 [avatar split clause 44,65977,65615]
65983. 28 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) [avatar definition]
65984. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) <- (~28) [avatar component clause 65983] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65985. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) <- (28) [avatar component clause 65983] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65986. 28 | ~14 [avatar split clause 43,65655,65983]
65987. 14 | ~28 [avatar split clause 42,65983,65655]
65989. 29 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) [avatar definition]
65990. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) <- (~29) [avatar component clause 65989] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65991. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) <- (29) [avatar component clause 65989] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65992. 29 | ~2 [avatar split clause 41,65607,65989]
65993. 2 | ~29 [avatar split clause 40,65989,65607]
65995. 30 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) [avatar definition]
65996. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) <- (~30) [avatar component clause 65995] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65997. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) <- (30) [avatar component clause 65995] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65998. 30 | ~12 [avatar split clause 39,65647,65995]
65999. 12 | ~30 [avatar split clause 38,65995,65647]
66001. 31 <=> v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) [avatar definition]
66002. ~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) <- (~31) [avatar component clause 66001] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66003. v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) <- (31) [avatar component clause 66001] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66004. 31 | ~7 [avatar split clause 37,65627,66001]
66005. 7 | ~31 [avatar split clause 36,66001,65627]
66007. 32 <=> v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) [avatar definition]
66008. ~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) <- (~32) [avatar component clause 66007] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66009. v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) <- (32) [avatar component clause 66007] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66010. 32 | ~9 [avatar split clause 35,65635,66007]
66011. 9 | ~32 [avatar split clause 34,66007,65635]
66013. 33 <=> v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) [avatar definition]
66014. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) <- (~33) [avatar component clause 66013] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66015. v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) <- (33) [avatar component clause 66013] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66016. 33 | ~13 [avatar split clause 33,65651,66013]
66017. 13 | ~33 [avatar split clause 32,66013,65651]
66019. 34 <=> v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) [avatar definition]
66020. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) <- (~34) [avatar component clause 66019] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66021. v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) <- (34) [avatar component clause 66019] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66022. 34 | ~11 [avatar split clause 31,65643,66019]
66023. 11 | ~34 [avatar split clause 30,66019,65643]
66025. 35 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) [avatar definition]
66026. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) <- (~35) [avatar component clause 66025] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66027. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) <- (35) [avatar component clause 66025] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66028. 35 | ~5 [avatar split clause 29,65619,66025]
66029. 5 | ~35 [avatar split clause 28,66025,65619]
66031. 36 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) [avatar definition]
66032. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) <- (~36) [avatar component clause 66031] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66033. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) <- (36) [avatar component clause 66031] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66034. 36 | ~15 [avatar split clause 27,65659,66031]
66035. 15 | ~36 [avatar split clause 26,66031,65659]
66037. 37 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) [avatar definition]
66038. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) <- (~37) [avatar component clause 66037] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66039. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) <- (37) [avatar component clause 66037] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66040. 37 | ~3 [avatar split clause 25,65611,66037]
66041. 3 | ~37 [avatar split clause 24,66037,65611]
66043. 38 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) [avatar definition]
66044. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) <- (~38) [avatar component clause 66043] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66045. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) <- (38) [avatar component clause 66043] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66046. 38 | ~1 [avatar split clause 23,65603,66043]
66047. 1 | ~38 [avatar split clause 22,66043,65603]
66048. v1_xboole_0(sK18) <- (~21) [resolution 65596,65944] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66050. 39 <=> v1_xboole_0(sK18) [avatar definition]
66052. v1_xboole_0(sK18) <- (39) [avatar component clause 66050] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66053. 39 | 21 [avatar split clause 66048,65943,66050]
66060. ~v1_xboole_0(sK19) <- (~25) [resolution 65598,65966] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66061. ~v1_xboole_0(k1_gate_1(sK19)) <- (~26) [resolution 65598,65972] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66062. ~v1_xboole_0(sK16) <- (~35) [resolution 65598,66026] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66063. ~v1_xboole_0(k1_gate_1(sK19)) <- (~28) [resolution 65598,65984] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66064. ~v1_xboole_0(sK19) <- (~29) [resolution 65598,65990] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66068. $false <- (22, ~25) [subsumption resolution 66060,65948] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66069. ~22 | 25 [avatar contradiction clause 66068]
66070. $false <- (19, ~35) [subsumption resolution 66062,65935] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66071. ~19 | 35 [avatar contradiction clause 66070]
66072. $false <- (22, ~29) [subsumption resolution 66064,65948] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66073. ~22 | 29 [avatar contradiction clause 66072]
66075. 40 <=> v1_xboole_0(k1_gate_1(sK19)) [avatar definition]
66076. v1_xboole_0(k1_gate_1(sK19)) <- (40) [avatar component clause 66075] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66077. ~v1_xboole_0(k1_gate_1(sK19)) <- (~40) [avatar component clause 66075] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66078. ~40 | 28 [avatar split clause 66063,65983,66075]
66079. ~40 | 26 [avatar split clause 66061,65971,66075]
66080. v1_xboole_0(sK19) <- (~40) [resolution 66077,65596] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66082. ~v1_xboole_0(sK21) <- (~26) [resolution 65599,65972] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66083. ~v1_xboole_0(k1_gate_1(sK21)) <- (~28) [resolution 65599,65984] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66087. ~18 | 26 [avatar split clause 66082,65971,65927]
66089. 41 <=> v1_xboole_0(k1_gate_1(sK21)) [avatar definition]
66090. v1_xboole_0(k1_gate_1(sK21)) <- (41) [avatar component clause 66089] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66091. ~v1_xboole_0(k1_gate_1(sK21)) <- (~41) [avatar component clause 66089] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66092. ~41 | 28 [avatar split clause 66083,65983,66089]
66094. ~v1_xboole_0(k1_gate_1(sK20)) <- (~26) [resolution 65600,65972] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66095. ~v1_xboole_0(sK20) <- (~28) [resolution 65600,65984] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66099. ~20 | 28 [avatar split clause 66095,65983,65937]
66101. 42 <=> v1_xboole_0(k1_gate_1(sK20)) [avatar definition]
66102. v1_xboole_0(k1_gate_1(sK20)) <- (42) [avatar component clause 66101] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66103. ~v1_xboole_0(k1_gate_1(sK20)) <- (~42) [avatar component clause 66101] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66104. ~42 | 26 [avatar split clause 66094,65971,66101]
66111. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (36) [resolution 65601,66033] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66112. v1_xboole_0(sK16) | v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (35) [resolution 65601,66027] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66113. v1_xboole_0(sK16) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (37) [resolution 65601,66039] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66114. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (30) [resolution 65601,65997] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66115. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (38) [resolution 65601,66045] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66116. v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (29) [resolution 65601,65991] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66119. v1_xboole_0(sK16) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (33) [resolution 65601,66015] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66121. v1_xboole_0(sK19) | v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (23) [resolution 65601,65955] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66123. v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(sK21) <- (25) [resolution 65601,65967] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66133. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (~19, 33) [subsumption resolution 66119,65934] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66141. 43 <=> v1_xboole_0(k1_gate_1(sK17)) [avatar definition]
66142. ~v1_xboole_0(k1_gate_1(sK17)) <- (~43) [avatar component clause 66141] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66143. v1_xboole_0(k1_gate_1(sK17)) <- (43) [avatar component clause 66141] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66145. 44 <=> v1_xboole_0(k1_gate_1(sK16)) [avatar definition]
66146. ~v1_xboole_0(k1_gate_1(sK16)) <- (~44) [avatar component clause 66145] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66147. v1_xboole_0(k1_gate_1(sK16)) <- (44) [avatar component clause 66145] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66148. 21 | 43 | 44 | ~38 [avatar split clause 66115,66043,66145,66141,65943]
66149. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (~19, 37) [subsumption resolution 66113,65934] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66150. 21 | 43 | 19 | ~37 [avatar split clause 66149,66037,65933,66141,65943]
66151. v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (~19, 35) [subsumption resolution 66112,65934] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66158. v1_xboole_0(sK19) | v1_xboole_0(sK21) <- (~20, 23) [subsumption resolution 66121,65939] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66161. 41 | 42 | 22 | ~29 [avatar split clause 66116,65989,65947,66101,66089]
66164. 22 | 40 [avatar split clause 66080,66075,65947]
66166. 39 | 43 | 19 | ~33 [avatar split clause 66133,66013,65933,66141,66050]
66170. 21 | 17 | 44 | ~36 [avatar split clause 66111,66031,66145,65923,65943]
66171. 21 | 17 | 19 | ~35 [avatar split clause 66151,66025,65933,65923,65943]
66174. 18 | 22 | 20 | ~23 [avatar split clause 66158,65953,65937,65947,65927]
66175. 18 | 42 | 22 | ~25 [avatar split clause 66123,65965,65947,66101,65927]
66176. 41 | 42 | 40 | ~30 [avatar split clause 66114,65995,66075,66101,66089]
66180. ~v1_xboole_0(sK21) <- (~24) [resolution 65960,65599] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66183. $false <- (18, ~24) [subsumption resolution 66180,65928] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66184. ~18 | 24 [avatar contradiction clause 66183]
66185. ~v1_xboole_0(sK19) <- (40) [resolution 66076,65597] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66187. ~v1_xboole_0(sK20) <- (~27) [resolution 65978,65600] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66188. ~v1_xboole_0(k1_gate_1(sK21)) <- (~27) [resolution 65978,65599] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66190. ~41 | 27 [avatar split clause 66188,65977,66089]
66191. v1_xboole_0(sK21) <- (~41) [resolution 66091,65596] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66194. ~v1_xboole_0(sK18) <- (~33) [resolution 66014,65599] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66197. $false <- (~33, 39) [subsumption resolution 66194,66052] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66198. 33 | ~39 [avatar contradiction clause 66197]
66201. v1_xboole_0(sK20) <- (~42) [resolution 66103,65596] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66205. 20 | 42 [avatar split clause 66201,66101,65937]
66206. ~20 | 27 [avatar split clause 66187,65977,65937]
66207. v1_xboole_0(sK19) | v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (27) [resolution 65979,65601] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66208. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(sK21) <- (26) [resolution 65973,65601] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66209. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (28) [resolution 65985,65601] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66215. ~v1_xboole_0(sK18) <- (~34) [resolution 66020,65599] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66218. $false <- (~34, 39) [subsumption resolution 66215,66052] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66219. 34 | ~39 [avatar contradiction clause 66218]
66224. ~v1_xboole_0(sK16) <- (~37) [resolution 66038,65598] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66226. $false <- (19, ~37) [subsumption resolution 66224,65935] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66227. ~19 | 37 [avatar contradiction clause 66226]
66231. ~v1_xboole_0(sK17) <- (~31) [resolution 66002,65600] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66232. ~v1_xboole_0(sK18) <- (~31) [resolution 66002,65599] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66233. ~v1_xboole_0(sK16) <- (~31) [resolution 66002,65598] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66237. $false <- (17, ~31) [subsumption resolution 66231,65925] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66238. ~17 | 31 [avatar contradiction clause 66237]
66239. $false <- (~31, 39) [subsumption resolution 66232,66052] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66240. 31 | ~39 [avatar contradiction clause 66239]
66241. $false <- (19, ~31) [subsumption resolution 66233,65935] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66242. ~19 | 31 [avatar contradiction clause 66241]
66246. ~v1_xboole_0(sK17) <- (~32) [resolution 66008,65600] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66247. ~v1_xboole_0(sK18) <- (~32) [resolution 66008,65599] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66251. $false <- (17, ~32) [subsumption resolution 66246,65925] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66252. ~17 | 32 [avatar contradiction clause 66251]
66253. $false <- (~32, 39) [subsumption resolution 66247,66052] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66254. 32 | ~39 [avatar contradiction clause 66253]
66257. ~v1_xboole_0(sK21) <- (~25) [resolution 65966,65599] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66262. ~18 | 25 [avatar split clause 66257,65965,65927]
66263. 18 | 41 [avatar split clause 66191,66089,65927]
66266. ~v1_xboole_0(sK16) <- (44) [resolution 66147,65597] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66268. $false <- (19, 44) [subsumption resolution 66266,65935] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66269. ~19 | ~44 [avatar contradiction clause 66268]
66270. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (32) [resolution 66009,65601] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66273. ~v1_xboole_0(sK17) <- (43) [resolution 66143,65597] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66275. $false <- (17, 43) [subsumption resolution 66273,65925] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66276. ~17 | ~43 [avatar contradiction clause 66275]
66279. ~v1_xboole_0(sK20) <- (~23) [resolution 65954,65600] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66280. ~v1_xboole_0(sK21) <- (~23) [resolution 65954,65599] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66284. $false <- (20, ~23) [subsumption resolution 66279,65938] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66285. ~20 | 23 [avatar contradiction clause 66284]
66286. $false <- (18, ~23) [subsumption resolution 66280,65928] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66287. ~18 | 23 [avatar contradiction clause 66286]
66290. ~v1_xboole_0(sK17) <- (~36) [resolution 66032,65600] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66294. $false <- (17, ~36) [subsumption resolution 66290,65925] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66295. ~17 | 36 [avatar contradiction clause 66294]
66296. ~22 | ~40 [avatar split clause 66185,66075,65947]
66300. ~v1_xboole_0(k1_gate_1(sK17)) <- (~34) [resolution 66020,65600] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66302. ~v1_xboole_0(k1_gate_1(sK16)) <- (~34) [resolution 66020,65598] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66303. ~43 | 34 [avatar split clause 66300,66019,66141]
66304. ~44 | 34 [avatar split clause 66302,66019,66145]
66309. ~v1_xboole_0(k1_gate_1(sK19)) <- (~30) [resolution 65996,65598] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66313. ~v1_xboole_0(k1_gate_1(sK18)) <- (~38) [resolution 66044,65599] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66316. $false <- (21, ~38) [subsumption resolution 66313,65945] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66317. ~21 | 38 [avatar contradiction clause 66316]
66321. ~v1_xboole_0(sK16) <- (~33) [resolution 66014,65598] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66323. $false <- (19, ~33) [subsumption resolution 66321,65935] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66324. ~19 | 33 [avatar contradiction clause 66323]
66325. ~40 | 30 [avatar split clause 66309,65995,66075]
66326. 39 | 17 | 44 | ~32 [avatar split clause 66270,66007,66145,65923,66050]
66327. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK21) <- (26, ~42) [subsumption resolution 66208,66103] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66328. 18 | 40 | ~26 | 42 [avatar split clause 66327,66101,65971,66075,65927]
66335. ~v1_xboole_0(sK21) <- (41) [resolution 66090,65597] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66338. ~v1_xboole_0(k1_gate_1(sK21)) <- (~29) [resolution 65990,65599] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66341. $false <- (~29, 41) [subsumption resolution 66338,66090] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66342. 29 | ~41 [avatar contradiction clause 66341]
66345. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (34) [resolution 66021,65601] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66346. v1_xboole_0(sK16) <- (~44) [resolution 66146,65596] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66360. ~v1_xboole_0(sK20) <- (~24) [resolution 65960,65600] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66362. ~v1_xboole_0(k1_gate_1(sK19)) <- (~24) [resolution 65960,65598] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66365. $false <- (20, ~24) [subsumption resolution 66360,65938] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66366. ~20 | 24 [avatar contradiction clause 66365]
66367. $false <- (~24, 40) [subsumption resolution 66362,66076] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66368. 24 | ~40 [avatar contradiction clause 66367]
66370. ~v1_xboole_0(k1_gate_1(sK17)) <- (~38) [resolution 66044,65600] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66373. ~43 | 38 [avatar split clause 66370,66043,66141]
66375. v1_xboole_0(sK17) <- (~43) [resolution 66142,65596] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66379. 17 | 43 [avatar split clause 66375,66141,65923]
66380. ~18 | ~41 [avatar split clause 66335,66089,65927]
66381. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (34, ~44) [subsumption resolution 66345,66146] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66382. 39 | 43 | ~34 | 44 [avatar split clause 66381,66145,66019,66141,66050]
66383. ~v1_xboole_0(sK18) <- (21) [resolution 65945,65597] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66384. ~39 | ~21 [avatar split clause 66383,65943,66050]
66388. ~v1_xboole_0(k1_gate_1(sK16)) <- (~32) [resolution 66008,65598] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66397. ~v1_xboole_0(k1_gate_1(sK21)) <- (~30) [resolution 65996,65599] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66400. $false <- (~30, 41) [subsumption resolution 66397,66090] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66401. 30 | ~41 [avatar contradiction clause 66400]
66404. ~v1_xboole_0(k1_gate_1(sK18)) <- (~36) [resolution 66032,65599] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66407. $false <- (21, ~36) [subsumption resolution 66404,65945] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66408. ~21 | 36 [avatar contradiction clause 66407]
66410. ~v1_xboole_0(sK20) <- (42) [resolution 66102,65597] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66414. 19 | 44 [avatar split clause 66346,66145,65933]
66415. ~20 | ~42 [avatar split clause 66410,66101,65937]
66416. ~44 | 32 [avatar split clause 66388,66007,66145]
66420. ~v1_xboole_0(sK19) <- (~23) [resolution 65954,65598] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66422. $false <- (22, ~23) [subsumption resolution 66420,65948] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66423. ~22 | 23 [avatar contradiction clause 66422]
66438. ~v1_xboole_0(k1_gate_1(sK18)) <- (~37) [resolution 66038,65599] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66444. $false <- (21, ~37) [subsumption resolution 66438,65945] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66445. ~21 | 37 [avatar contradiction clause 66444]
66448. ~v1_xboole_0(k1_gate_1(sK18)) <- (~35) [resolution 66026,65599] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66451. $false <- (21, ~35) [subsumption resolution 66448,65945] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66452. ~21 | 35 [avatar contradiction clause 66451]
66464. v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK21)) <- (~20, 27) [subsumption resolution 66207,65939] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66465. 41 | 22 | 20 | ~27 [avatar split clause 66464,65977,65937,65947,66089]
66466. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK21)) <- (~20, 28) [subsumption resolution 66209,65939] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66467. 41 | 40 | 20 | ~28 [avatar split clause 66466,65983,65937,66075,66089]
66473. v1_xboole_0(sK16) | v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (31) [resolution 66003,65601] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66475. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (24) [resolution 65961,65601] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66478. ~v1_xboole_0(k1_gate_1(sK20)) <- (~25) [resolution 65966,65600] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66482. $false <- (~25, 42) [subsumption resolution 66478,66102] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66483. 25 | ~42 [avatar contradiction clause 66482]
66492. ~v1_xboole_0(sK17) <- (~35) [resolution 66026,65600] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66496. $false <- (17, ~35) [subsumption resolution 66492,65925] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66497. ~17 | 35 [avatar contradiction clause 66496]
66499. ~v1_xboole_0(k1_gate_1(sK17)) <- (~37) [resolution 66038,65600] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66502. ~43 | 37 [avatar split clause 66499,66037,66141]
66508. ~v1_xboole_0(k1_gate_1(sK16)) <- (~38) [resolution 66044,65598] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66510. $false <- (~38, 44) [subsumption resolution 66508,66147] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66511. 38 | ~44 [avatar contradiction clause 66510]
66513. ~v1_xboole_0(k1_gate_1(sK20)) <- (~29) [resolution 65990,65600] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66517. $false <- (~29, 42) [subsumption resolution 66513,66102] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66518. 29 | ~42 [avatar contradiction clause 66517]
66532. ~v1_xboole_0(k1_gate_1(sK16)) <- (~36) [resolution 66032,65598] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66534. $false <- (~36, 44) [subsumption resolution 66532,66147] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66535. 36 | ~44 [avatar contradiction clause 66534]
66536. v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (~19, 31) [subsumption resolution 66473,65934] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66537. 39 | 17 | 19 | ~31 [avatar split clause 66536,66001,65933,65923,66050]
66538. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK21) <- (~20, 24) [subsumption resolution 66475,65939] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66539. 18 | 40 | 20 | ~24 [avatar split clause 66538,65959,65937,66075,65927]
66546. ~v1_xboole_0(sK19) <- (~27) [resolution 65978,65598] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66548. $false <- (22, ~27) [subsumption resolution 66546,65948] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66549. ~22 | 27 [avatar contradiction clause 66548]
66552. ~v1_xboole_0(k1_gate_1(sK17)) <- (~33) [resolution 66014,65600] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66555. ~43 | 33 [avatar split clause 66552,66013,66141]
66564. ~v1_xboole_0(k1_gate_1(sK20)) <- (~30) [resolution 65996,65600] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66568. $false <- (~30, 42) [subsumption resolution 66564,66102] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66569. 30 | ~42 [avatar contradiction clause 66568]
66570. $false [avatar sat refutation 65793,65857,65889,65905,65913,65917,65919,65920,65930,65931,65940,65941,65950,65951,65956,65957,65962,65963,65968,65969,65974,65975,65980,65981,65986,65987,65992,65993,65998,65999,66004,66005,66010,66011,66016,66017,66022,66023,66028,66029,66034,66035,66040,66041,66046,66047,66053,66069,66071,66073,66078,66079,66087,66092,66099,66104,66148,66150,66161,66164,66166,66170,66171,66174,66175,66176,66184,66190,66198,66205,66206,66219,66227,66238,66240,66242,66252,66254,66262,66263,66269,66276,66285,66287,66295,66296,66303,66304,66317,66324,66325,66326,66328,66342,66366,66368,66373,66379,66380,66382,66384,66401,66408,66414,66415,66416,66423,66445,66452,66465,66467,66483,66497,66502,66511,66518,66535,66537,66539,66549,66555,66569] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:5504,thDist:-5504}
% SZS output end Proof for gate_3__t3_gate_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 22
% 
% >>> Saturation
% Initial clauses: 65580
% Generated clauses: 66333
% Active clauses: 297
% Passive clauses: 339
% Final active clauses: 49
% Final passive clauses: 1
% Main loop iterations started: 297
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 71
% 
% >>> Deletion Inferences
% Simple tautologies: 65284
% Forward subsumptions: 181
% 
% >>> Generating Inferences
% Binary resolution: 235
% Unit resulting resolution: 45
% 
% >>> AVATAR
% Split clauses: 318
% Split components: 4228
% Sat splitting refutations: 90
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 443
% SAT solver binary clauses: 153
% 
% Memory used [KB]: 35692
% Time elapsed: 0.123 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.047 s
% preprocessing: 0.012 s
% SAT solver time: 0.001 s
% term sharing: 0.001 s
% congruence closure: 0.001 s
% backward subsumption index maintenance: 0.001 s
% resolution: 0.001 s
% unit resulting resolution: 0.003 s ( own 0.002 s ) 
% other: 0.123 s ( own 0.057 s ) 

