i: [1,53,0,3,1,1,0,0] t2_arytm_1
a: [3,75,0,3,1,53]
s: 75
i: [1,52,0,3,1,1,0,0] t2_arytm_1
a: [3,80,0,3,1,52]
s: 80
i: [1,50,0,3,1,1,0,0] t2_arytm_1
a: [3,90,0,3,1,50]
s: 90
i: [1,49,0,3,1,1,0,0] t2_arytm_1
a: [3,95,0,3,1,49]
s: 95
i: [1,51,0,5,1,1,0,0] t2_arytm_1
a: [3,85,0,5,1,51]
s: 85
i: [1,57,0,2,1,0,0,255] dt_k12_arytm_3
a: [3,120,0,2,1,57]
s: 120
i: [1,54,0,3,1,0,0,255] redefinition_k12_arytm_3
a: [3,100,0,3,1,54]
s: 100
i: [1,55,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,105,0,3,1,55]
s: 105
i: [1,56,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,110,0,3,1,56]
s: 110
i: [1,58,0,3,1,0,0,255] dt_k12_arytm_3
a: [3,115,0,3,1,58]
s: 115
i: [1,70,0,11,3,0,0,255] t15_arytm_2
a: [3,125,0,3,1,70]
s: 125
i: [1,71,0,11,3,0,0,1] t4_arytm_2
a: [3,133,0,3,1,71]
s: 133
i: [1,66,0,4,1,0,0,1] existence_m1_subset_1
s: 66
i: [1,64,0,6,2,0,0,255] t1_subset
s: 64
i: [1,67,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 67
i: [1,65,0,8,3,0,0,255] t2_subset
s: 65
d: [2,138,1,5,2,1,40,65,90]
a: [3,150,1,2,1,138]
s: 150
i: [1,59,0,10,3,0,0,255] t14_arytm_2
s: 59
d: [2,140,1,5,2,1,40,65,110]
a: [3,164,1,2,1,140]
s: 164
i: [1,68,0,11,3,0,0,1] dt_k8_arytm_2
s: 68
d: [2,173,1,9,3,1,44,68,85]
d: [2,174,1,6,2,2,21,173,133]
d: [2,175,1,3,1,3,21,174,90]
d: [2,176,1,0,0,4,21,175,95]
e: 176
r: 133
a: [3,132,0,3,1,71]
s: 132
d: [2,142,1,6,2,0,40,65,66]
s: 142
a: [3,136,0,8,2,71]
s: 136
d: [2,182,2,6,2,0,40,142,67]
s: 182
i: [1,60,0,12,3,0,0,255] t14_arytm_2
s: 60
i: [1,69,0,13,3,0,0,1] commutativity_k8_arytm_2
s: 69
d: [2,187,1,16,5,0,44,68,60]
d: [2,189,1,13,4,0,18,187]
d: [2,190,1,10,3,1,21,189,125]
d: [2,191,1,6,2,1,21,190,59]
s: 191
r: 60
r: 59
d: [2,208,2,0,0,3,52,75,90,191]
e: 208
d: [2,209,2,0,0,3,52,80,95,191]
e: 209
d: [2,212,2,3,1,2,40,191,95]
d: [2,217,2,0,0,3,21,212,80]
e: 217
d: [2,213,2,3,1,2,40,191,90]
d: [2,219,2,0,0,3,21,213,75]
e: 219
r: 191
r: 125
a: [3,124,0,3,1,70]
s: 124
a: [3,128,0,8,2,70]
s: 128
s: 59
s: 60
d: [2,242,1,14,4,1,44,60,128]
d: [2,247,1,11,3,2,21,242,124]
a: [3,251,1,4,1,247]
s: 251
i: [1,63,0,20,4,0,0,1] t12_arytm_2
s: 63
d: [2,272,1,18,4,1,44,63,85]
d: [2,295,1,15,3,2,21,272,90]
d: [2,296,1,12,2,3,21,295,95]
s: 296
a: [3,255,1,3,1,247]
s: 255
r: 60
r: 100
r: 115
r: 120
r: 124
r: 128
r: 251
d: [2,330,0,2,1,2,23,120,255]
a: [3,342,0,2,1,330]
s: 342
d: [2,307,2,18,4,3,44,296,60]
d: [2,314,2,15,3,4,21,307,75]
d: [2,315,2,12,2,4,21,314,90]
a: [3,319,2,4,1,315]
s: 319
d: [2,349,3,0,0,3,52,90,75,319,59]
e: 349
d: [2,350,3,6,2,1,40,319,59]
d: [2,352,3,3,1,2,21,350,75]
d: [2,353,3,0,0,3,21,352,90]
e: 353
r: 319
d: [2,328,0,3,1,2,23,100,255]
a: [3,337,0,3,1,328]
s: 337
a: [3,318,2,4,1,315]
s: 318
d: [2,332,0,8,2,2,23,128,255]
s: 332
d: [2,327,0,12,3,1,23,60,255]
s: 327
a: [3,323,2,8,1,315]
d: [2,355,2,8,1,2,22,323,255]
a: [3,359,2,8,1,355]
s: 359
d: [2,333,1,4,1,2,23,251,255]
a: [3,347,1,4,1,333]
s: 347
d: [2,401,3,9,2,2,44,359,136]
d: [2,412,3,5,1,3,21,401,318]
a: [3,416,3,5,1,412]
s: 416
r: 359
d: [2,365,1,9,3,3,44,136,332]
d: [2,374,1,6,2,3,18,365]
s: 374
r: 59
r: 327
d: [2,446,2,0,0,5,52,80,95,374]
e: 446
d: [2,447,2,0,0,5,52,75,90,374]
e: 447
d: [2,450,2,3,1,4,40,374,95]
d: [2,456,2,0,0,5,21,450,80]
e: 456
d: [2,451,2,3,1,4,40,374,90]
d: [2,458,2,0,0,5,21,451,75]
e: 458
r: 374
r: 332
r: 255
r: 337
s: 120
d: [2,405,3,12,3,1,44,68,359]
d: [2,424,3,8,2,2,21,405,318]
d: [2,425,3,5,1,3,21,424,132]
a: [3,429,3,5,1,425]
s: 429
s: 100
d: [2,432,2,6,1,2,23,359,416]
a: [3,442,2,6,1,432]
s: 442
s: 115
s: 124
d: [2,400,3,10,2,3,44,359,332]
d: [2,406,3,6,1,4,21,400,318]
a: [3,410,3,6,1,406]
s: 410
s: 128
d: [2,468,3,4,1,2,22,410,416]
a: [3,471,3,4,1,468]
s: 471
s: 59
s: 60
r: 442
a: [3,441,2,6,1,432]
s: 441
d: [2,526,3,0,0,3,52,318,441,136]
e: 526
d: [2,527,3,7,2,2,44,441,136]
d: [2,528,3,4,1,2,19,527]
d: [2,530,3,0,0,3,21,528,318]
e: 530
r: 441
a: [3,254,1,3,1,247]
s: 254
d: [2,503,2,18,4,3,44,296,60]
d: [2,514,2,14,3,4,21,503,318]
d: [2,515,2,11,2,5,21,514,75]
d: [2,516,2,8,1,5,21,515,90]
d: [2,517,2,5,1,6,22,516,442]
a: [3,521,2,5,1,517]
s: 521
a: [3,250,1,4,1,247]
s: 250
a: [3,259,1,4,1,247]
s: 259
r: 250
d: [2,534,3,11,3,1,44,521,69]
d: [2,540,3,8,2,2,21,534,95]
d: [2,541,3,5,1,3,21,540,124]
a: [3,545,3,5,1,541]
s: 545
d: [2,562,4,6,2,2,44,545,128]
d: [2,566,4,3,1,3,21,562,80]
d: [2,567,4,0,0,4,21,566,95]
e: 567
d: [2,563,4,6,2,2,44,128,545]
d: [2,569,4,3,1,3,21,563,80]
d: [2,570,4,0,0,4,21,569,95]
e: 570
d: [2,178,1,6,2,2,21,173,90]
d: [2,179,1,3,1,3,21,178,95]
d: [2,538,3,11,3,1,44,69,521]
d: [2,555,3,8,2,2,21,538,124]
d: [2,556,3,5,1,3,21,555,95]
f: 570,556,530,517,353,217,179,71,70,49,50,51,52,53
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t2_arytm_1
% SZS output start Proof for arytm_1__t2_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ~(k11_arytm_3 != X1 & k11_arytm_3 != X0 & k8_arytm_2(X0,X1) = k11_arytm_3))) [input t2_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ~(k11_arytm_3 != X1 & k11_arytm_3 != X0 & k8_arytm_2(X0,X1) = k11_arytm_3))) [negated conjecture 1]
4. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => k8_arytm_2(X0,X1) = k8_arytm_2(X1,X0)) [input commutativity_k8_arytm_2]
11. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => m1_subset_1(k8_arytm_2(X0,X1),k2_arytm_2)) [input dt_k8_arytm_2]
16. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => k8_arytm_2(X0,k8_arytm_2(X1,X2)) = k8_arytm_2(k8_arytm_2(X0,X1),X2)))) [input t12_arytm_2]
17. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ~(! [X1] : (m1_subset_1(X1,k2_arytm_2) => k8_arytm_2(X0,X1) != k12_arytm_3) & k11_arytm_3 != X0)) [input t14_arytm_2]
18. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k12_arytm_3 = X0 => k8_arytm_2(X0,X1) = X1))) [input t15_arytm_2]
21. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => k8_arytm_2(X0,X1) = k11_arytm_3))) [input t4_arytm_2]
25. ? [X0] : (? [X1] : ((k11_arytm_3 != X1 & k11_arytm_3 != X0 & k8_arytm_2(X0,X1) = k11_arytm_3) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
26. ? [X0] : (? [X1] : (k11_arytm_3 != X1 & k11_arytm_3 != X0 & k8_arytm_2(X0,X1) = k11_arytm_3 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 25]
27. ! [X0] : ((? [X1] : (k8_arytm_2(X0,X1) = k12_arytm_3 & m1_subset_1(X1,k2_arytm_2)) | k11_arytm_3 = X0) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 17]
28. ! [X0] : (? [X1] : (k8_arytm_2(X0,X1) = k12_arytm_3 & m1_subset_1(X1,k2_arytm_2)) | k11_arytm_3 = X0 | ~m1_subset_1(X0,k2_arytm_2)) [flattening 27]
29. ! [X0] : (! [X1] : ((k8_arytm_2(X0,X1) = X1 | k12_arytm_3 != X0) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 18]
30. ! [X0] : (! [X1] : (k8_arytm_2(X0,X1) = X1 | k12_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 29]
31. ! [X0] : (! [X1] : ((k8_arytm_2(X0,X1) = k11_arytm_3 | k11_arytm_3 != X0) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 21]
32. ! [X0] : (! [X1] : (k8_arytm_2(X0,X1) = k11_arytm_3 | k11_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 31]
33. ! [X0] : (! [X1] : (! [X2] : (k8_arytm_2(X0,k8_arytm_2(X1,X2)) = k8_arytm_2(k8_arytm_2(X0,X1),X2) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 16]
38. ! [X0,X1] : (m1_subset_1(k8_arytm_2(X0,X1),k2_arytm_2) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 11]
39. ! [X0,X1] : (m1_subset_1(k8_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 38]
40. ! [X0,X1] : (k8_arytm_2(X0,X1) = k8_arytm_2(X1,X0) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 4]
41. ! [X0,X1] : (k8_arytm_2(X0,X1) = k8_arytm_2(X1,X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 40]
42. ? [X0] : (? [X1] : (k11_arytm_3 != X1 & k11_arytm_3 != X0 & k8_arytm_2(X0,X1) = k11_arytm_3 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (k11_arytm_3 != X1 & k11_arytm_3 != sK0 & k11_arytm_3 = k8_arytm_2(sK0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
43. ? [X1] : (k11_arytm_3 != X1 & k11_arytm_3 != sK0 & k11_arytm_3 = k8_arytm_2(sK0,X1) & m1_subset_1(X1,k2_arytm_2)) => (k11_arytm_3 != sK1 & k11_arytm_3 != sK0 & k11_arytm_3 = k8_arytm_2(sK0,sK1) & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
44. (k11_arytm_3 != sK1 & k11_arytm_3 != sK0 & k11_arytm_3 = k8_arytm_2(sK0,sK1) & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 26,43,42]
45. ! [X0] : (? [X1] : (k8_arytm_2(X0,X1) = k12_arytm_3 & m1_subset_1(X1,k2_arytm_2)) => (k12_arytm_3 = k8_arytm_2(X0,sK2(X0)) & m1_subset_1(sK2(X0),k2_arytm_2))) [choice axiom]
46. ! [X0] : ((k12_arytm_3 = k8_arytm_2(X0,sK2(X0)) & m1_subset_1(sK2(X0),k2_arytm_2)) | k11_arytm_3 = X0 | ~m1_subset_1(X0,k2_arytm_2)) [skolemisation 28,45]
49. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. k11_arytm_3 = k8_arytm_2(sK0,sK1) [cnf transformation 44] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. k11_arytm_3 != sK0 [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
53. k11_arytm_3 != sK1 [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
59. m1_subset_1(sK2(X0),k2_arytm_2) | k11_arytm_3 = X0 | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 46] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. k12_arytm_3 = k8_arytm_2(X0,sK2(X0)) | k11_arytm_3 = X0 | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 46] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
61. k8_arytm_2(X0,X1) = X1 | k12_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 30] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
62. k8_arytm_2(X0,X1) = k11_arytm_3 | k11_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 32] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:1,thDist:-1}
63. k8_arytm_2(X0,k8_arytm_2(X1,X2)) = k8_arytm_2(k8_arytm_2(X0,X1),X2) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 33] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. m1_subset_1(k8_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 39] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. k8_arytm_2(X0,X1) = k8_arytm_2(X1,X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 41] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. k8_arytm_2(k12_arytm_3,X1) = X1 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k12_arytm_3,k2_arytm_2) [equality resolution 61] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
71. k11_arytm_3 = k8_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) [equality resolution 62] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:1,thDist:-1}
73. 1 <=> k11_arytm_3 = sK1 [avatar definition]
75. k11_arytm_3 != sK1 <- (~1) [avatar component clause 73] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
76. ~1 [avatar split clause 53,73]
78. 2 <=> k11_arytm_3 = sK0 [avatar definition]
80. k11_arytm_3 != sK0 <- (~2) [avatar component clause 78] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
81. ~2 [avatar split clause 52,78]
83. 3 <=> k11_arytm_3 = k8_arytm_2(sK0,sK1) [avatar definition]
85. k11_arytm_3 = k8_arytm_2(sK0,sK1) <- (3) [avatar component clause 83] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. 3 [avatar split clause 51,83]
88. 4 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
90. m1_subset_1(sK1,k2_arytm_2) <- (4) [avatar component clause 88] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. 4 [avatar split clause 50,88]
93. 5 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
95. m1_subset_1(sK0,k2_arytm_2) <- (5) [avatar component clause 93] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. 5 [avatar split clause 49,93]
123. 11 <=> m1_subset_1(k12_arytm_3,k2_arytm_2) [avatar definition]
124. m1_subset_1(k12_arytm_3,k2_arytm_2) <- (11) [avatar component clause 123] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
125. ~m1_subset_1(k12_arytm_3,k2_arytm_2) <- (~11) [avatar component clause 123] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
127. 12 <=> ! [X1] : (k8_arytm_2(k12_arytm_3,X1) = X1 | ~m1_subset_1(X1,k2_arytm_2)) [avatar definition]
128. k8_arytm_2(k12_arytm_3,X1) = X1 | ~m1_subset_1(X1,k2_arytm_2) <- (12) [avatar component clause 127] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
129. ~11 | 12 [avatar split clause 70,127,123]
131. 13 <=> m1_subset_1(k11_arytm_3,k2_arytm_2) [avatar definition]
135. 14 <=> ! [X1] : (k11_arytm_3 = k8_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2)) [avatar definition]
136. k11_arytm_3 = k8_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2) <- (14) [avatar component clause 135] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
137. ~13 | 14 [avatar split clause 71,135,131]
173. m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (3) [superposition 68,85] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
178. m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (3, 4) [subsumption resolution 173,90] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
179. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (3, 4, 5) [subsumption resolution 178,95] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
180. 13 | ~3 | ~4 | ~5 [avatar split clause 179,93,88,83,131]
187. m1_subset_1(k12_arytm_3,k2_arytm_2) | ~m1_subset_1(sK2(X0),k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) | k11_arytm_3 = X0 | ~m1_subset_1(X0,k2_arytm_2) [superposition 68,60] {a:1,w:16,wCS:6400,sine:1,thAx:0,allAx:3,thDist:-3}
189. m1_subset_1(k12_arytm_3,k2_arytm_2) | ~m1_subset_1(sK2(X0),k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) | k11_arytm_3 = X0 [duplicate literal removal 187] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:3,thDist:-3}
190. ~m1_subset_1(sK2(X0),k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) | k11_arytm_3 = X0 <- (~11) [subsumption resolution 189,125] {a:1,w:10,wCS:4000,sine:1,thAx:0,allAx:3,thDist:-3}
191. ~m1_subset_1(X0,k2_arytm_2) | k11_arytm_3 = X0 <- (~11) [subsumption resolution 190,59] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
212. k11_arytm_3 = sK0 <- (5, ~11) [resolution 191,95] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
217. $false <- (~2, 5, ~11) [subsumption resolution 212,80] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
218. 2 | ~5 | 11 [avatar contradiction clause 217]
272. k8_arytm_2(sK0,k8_arytm_2(sK1,X16)) = k8_arytm_2(k11_arytm_3,X16) | ~m1_subset_1(X16,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (3) [superposition 63,85] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
295. k8_arytm_2(sK0,k8_arytm_2(sK1,X16)) = k8_arytm_2(k11_arytm_3,X16) | ~m1_subset_1(X16,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (3, 4) [subsumption resolution 272,90] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
296. k8_arytm_2(sK0,k8_arytm_2(sK1,X16)) = k8_arytm_2(k11_arytm_3,X16) | ~m1_subset_1(X16,k2_arytm_2) <- (3, 4, 5) [subsumption resolution 295,95] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
317. 25 <=> m1_subset_1(sK2(sK1),k2_arytm_2) [avatar definition]
318. m1_subset_1(sK2(sK1),k2_arytm_2) <- (25) [avatar component clause 317] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
319. ~m1_subset_1(sK2(sK1),k2_arytm_2) <- (~25) [avatar component clause 317] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
350. k11_arytm_3 = sK1 | ~m1_subset_1(sK1,k2_arytm_2) <- (~25) [resolution 319,59] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
352. ~m1_subset_1(sK1,k2_arytm_2) <- (~1, ~25) [subsumption resolution 350,75] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
353. $false <- (~1, 4, ~25) [subsumption resolution 352,90] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
354. 1 | ~4 | 25 [avatar contradiction clause 353]
440. 34 <=> k11_arytm_3 = k8_arytm_2(k11_arytm_3,sK2(sK1)) [avatar definition]
441. k11_arytm_3 != k8_arytm_2(k11_arytm_3,sK2(sK1)) <- (~34) [avatar component clause 440] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
442. k11_arytm_3 = k8_arytm_2(k11_arytm_3,sK2(sK1)) <- (34) [avatar component clause 440] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
503. k8_arytm_2(k11_arytm_3,sK2(sK1)) = k8_arytm_2(sK0,k12_arytm_3) | ~m1_subset_1(sK2(sK1),k2_arytm_2) | k11_arytm_3 = sK1 | ~m1_subset_1(sK1,k2_arytm_2) <- (3, 4, 5) [superposition 296,60] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
514. k8_arytm_2(k11_arytm_3,sK2(sK1)) = k8_arytm_2(sK0,k12_arytm_3) | k11_arytm_3 = sK1 | ~m1_subset_1(sK1,k2_arytm_2) <- (3, 4, 5, 25) [subsumption resolution 503,318] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
515. k8_arytm_2(k11_arytm_3,sK2(sK1)) = k8_arytm_2(sK0,k12_arytm_3) | ~m1_subset_1(sK1,k2_arytm_2) <- (~1, 3, 4, 5, 25) [subsumption resolution 514,75] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
516. k8_arytm_2(k11_arytm_3,sK2(sK1)) = k8_arytm_2(sK0,k12_arytm_3) <- (~1, 3, 4, 5, 25) [subsumption resolution 515,90] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
517. k11_arytm_3 = k8_arytm_2(sK0,k12_arytm_3) <- (~1, 3, 4, 5, 25, 34) [forward demodulation 516,442] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
519. 36 <=> k11_arytm_3 = k8_arytm_2(sK0,k12_arytm_3) [avatar definition]
521. k11_arytm_3 = k8_arytm_2(sK0,k12_arytm_3) <- (36) [avatar component clause 519] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
522. 36 | 1 | ~3 | ~4 | ~5 | ~25 | ~34 [avatar split clause 517,440,317,93,88,83,73,519]
527. k11_arytm_3 != k11_arytm_3 | ~m1_subset_1(sK2(sK1),k2_arytm_2) <- (14, ~34) [superposition 441,136] {a:3,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
528. ~m1_subset_1(sK2(sK1),k2_arytm_2) <- (14, ~34) [trivial inequality removal 527] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
530. $false <- (14, 25, ~34) [subsumption resolution 528,318] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
531. ~14 | ~25 | 34 [avatar contradiction clause 530]
538. k11_arytm_3 = k8_arytm_2(k12_arytm_3,sK0) | ~m1_subset_1(k12_arytm_3,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (36) [superposition 69,521] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
543. 37 <=> k11_arytm_3 = k8_arytm_2(k12_arytm_3,sK0) [avatar definition]
545. k11_arytm_3 = k8_arytm_2(k12_arytm_3,sK0) <- (37) [avatar component clause 543] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
555. k11_arytm_3 = k8_arytm_2(k12_arytm_3,sK0) | ~m1_subset_1(sK0,k2_arytm_2) <- (11, 36) [subsumption resolution 538,124] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
556. k11_arytm_3 = k8_arytm_2(k12_arytm_3,sK0) <- (5, 11, 36) [subsumption resolution 555,95] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
557. 37 | ~5 | ~11 | ~36 [avatar split clause 556,519,123,93,543]
563. k11_arytm_3 = sK0 | ~m1_subset_1(sK0,k2_arytm_2) <- (12, 37) [superposition 128,545] {a:4,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
569. ~m1_subset_1(sK0,k2_arytm_2) <- (~2, 12, 37) [subsumption resolution 563,80] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
570. $false <- (~2, 5, 12, 37) [subsumption resolution 569,95] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
571. 2 | ~5 | ~12 | ~37 [avatar contradiction clause 570]
574. $false [avatar sat refutation 76,81,86,91,96,129,137,180,218,354,522,531,557,571] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:453,thDist:-453}
% SZS output end Proof for arytm_1__t2_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 364
% Active clauses: 60
% Passive clauses: 110
% Final active clauses: 37
% Final passive clauses: 37
% Main loop iterations started: 62
% 
% >>> Simplifying Inferences
% Duplicate literals: 49
% Trivial inequalities: 1
% Fw subsumption resolutions: 99
% Fw demodulations: 7
% Bw demodulations: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Equational tautologies: 4
% Forward subsumptions: 27
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 32
% Unit resulting resolution: 6
% Forward superposition: 34
% Backward superposition: 54
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 28
% Sat splitting refutations: 15
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 63
% SAT solver unit clauses: 10
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% forward demodulation: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

