i: [1,38,0,18,2,1,0,0] t2_gate_2
a: [3,50,0,7,1,38]
s: 50
a: [3,55,0,11,1,38]
s: 55
i: [1,39,0,5,2,0,0,255] t6_boole
s: 39
i: [1,42,0,6,2,0,0,1] fc4_gate_1
s: 42
i: [1,40,0,7,3,0,0,255] t8_boole
s: 40
i: [1,41,0,7,1,0,0,255] commutativity_k2_gate_1
s: 41
r: 50
r: 55
d: [2,59,1,5,1,1,40,42,50]
a: [3,63,1,5,1,59]
s: 63
d: [2,65,0,7,1,1,23,50,41]
a: [3,72,0,7,1,65]
s: 72
d: [2,90,1,2,1,1,40,72,42]
a: [3,94,1,2,1,90]
s: 94
d: [2,66,0,11,1,1,23,55,41]
d: [2,74,0,11,1,1,22,66,41]
a: [3,78,0,11,1,74]
s: 78
i: [1,44,0,7,2,0,0,1] fc18_gate_1
s: 44
d: [2,96,2,2,1,1,40,44,63]
a: [3,100,2,2,1,96]
s: 100
i: [1,45,0,7,2,0,0,1] fc20_gate_1
s: 45
d: [2,102,2,2,1,1,40,45,63]
a: [3,106,2,2,1,102]
s: 106
i: [1,46,0,7,2,0,0,1] fc19_gate_1
s: 46
i: [1,43,0,8,3,0,0,1] fc3_gate_1
s: 43
d: [2,108,2,2,1,1,40,46,63]
a: [3,112,2,2,1,108]
s: 112
i: [1,47,0,11,4,0,0,1] fc17_gate_1
s: 47
d: [2,117,1,12,3,1,40,47,78]
a: [3,124,1,4,1,117]
s: 124
d: [2,134,3,0,0,3,52,100,94,124,43]
e: 134
d: [2,135,2,4,2,1,40,124,43]
d: [2,137,2,2,1,2,21,135,94]
d: [2,138,2,0,0,3,21,137,100]
e: 138
r: 63
r: 72
r: 78
r: 100
a: [3,71,0,7,1,65]
s: 71
a: [3,77,0,11,1,74]
s: 77
d: [2,149,2,0,0,2,52,124,77,44]
e: 149
d: [2,152,1,4,1,1,40,77,44]
d: [2,154,1,0,0,2,21,152,124]
e: 154
r: 124
a: [3,99,2,2,1,96]
s: 99
a: [3,123,1,4,1,117]
s: 123
d: [2,159,3,0,0,2,52,99,123,42]
e: 159
d: [2,160,2,2,1,1,40,123,42]
d: [2,162,2,0,0,2,21,160,99]
e: 162
r: 99
s: 100
a: [3,127,1,4,1,117]
s: 127
a: [3,131,1,4,1,117]
s: 131
a: [3,62,1,5,1,59]
s: 62
d: [2,167,3,0,0,4,52,100,112,106,62,47]
e: 167
d: [2,168,2,6,3,1,40,62,47]
d: [2,170,2,4,2,2,21,168,106]
d: [2,171,2,2,1,3,21,170,112]
d: [2,172,2,0,0,4,21,171,100]
e: 172
r: 106
d: [2,67,1,6,2,0,44,42,41]
s: 67
r: 112
r: 62
r: 94
r: 71
r: 77
r: 131
s: 94
s: 72
s: 78
s: 106
a: [3,132,1,4,1,117]
s: 132
d: [2,203,3,0,0,3,52,94,106,132,43]
e: 203
d: [2,204,2,4,2,1,40,132,43]
d: [2,206,2,2,1,2,21,204,106]
d: [2,207,2,0,0,3,21,206,94]
e: 207
r: 127
r: 132
s: 112
a: [3,128,1,4,1,117]
s: 128
d: [2,211,3,0,0,3,52,94,112,128,43]
e: 211
d: [2,212,2,4,2,1,40,128,43]
d: [2,214,2,2,1,2,21,212,112]
d: [2,215,2,0,0,3,21,214,94]
e: 215
i: [1,35,0,18,2,1,0,0] t2_gate_2
a: [3,54,0,11,1,38]
d: [2,141,0,11,1,1,22,54,41]
d: [2,142,0,11,1,1,22,141,41]
a: [3,51,0,7,1,38]
d: [2,144,0,7,1,1,22,51,41]
d: [2,151,1,4,1,1,40,77,45]
d: [2,150,1,4,1,1,40,77,46]
d: [2,180,2,2,1,1,40,67,131]
d: [2,179,2,2,1,1,40,67,127]
d: [2,185,2,4,2,2,21,168,100]
d: [2,165,2,2,1,1,40,127,42]
d: [2,146,1,7,2,1,40,71,43]
d: [2,191,0,11,1,1,22,55,41]
d: [2,192,0,11,1,1,22,191,41]
d: [2,194,0,7,1,1,22,50,41]
d: [2,199,1,12,3,1,40,78,47]
d: [2,209,1,8,2,2,21,199,123]
f: 215,209,207,194,192,146,165,185,179,180,160,150,151,152,144,142,137,108,102,96,90,59,35,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_2__t2_gate_2
% SZS output start Proof for gate_2__t2_gate_2
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (~(v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))) & ~v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1))) & ~(v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))))) [input t2_gate_2]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (~(v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))) & ~v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1))) & ~(v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))))) [negated conjecture 1]
4. ! [X0,X1] : k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [input commutativity_k2_gate_1]
8. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc17_gate_1]
9. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc18_gate_1]
10. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X0,X2))) [input fc19_gate_1]
11. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X2,X0))) [input fc20_gate_1]
12. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_gate_1(X0,X1))) [input fc3_gate_1]
13. ! [X0,X1] : (v1_xboole_0(X0) => v1_xboole_0(k2_gate_1(X0,X1))) [input fc4_gate_1]
17. ~! [X0,X1,X2,X3] : (~(v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))) & ~v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1))) & ~(v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))))) [flattening 2]
22. ? [X0,X1,X2,X3] : ((v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))) & ~v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1))) | (v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))))) [ennf transformation 17]
25. ! [X0,X1] : (v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 13]
26. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 12]
27. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 26]
28. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0)) [ennf transformation 9]
29. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0)) [ennf transformation 11]
30. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0)) [ennf transformation 10]
31. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 8]
32. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 31]
33. ? [X0,X1,X2,X3] : ((v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))) & ~v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1))) | (v1_xboole_0(k2_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k8_gate_1(k2_gate_1(X0,X1),k2_gate_1(X2,X1),k2_gate_1(X3,X1))))) => ((v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK3,sK1))) & ~v1_xboole_0(k2_gate_1(k8_gate_1(sK0,sK2,sK3),sK1))) | (v1_xboole_0(k2_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) & ~v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK3,sK1))))) [choice axiom]
34. (v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK3,sK1))) & ~v1_xboole_0(k2_gate_1(k8_gate_1(sK0,sK2,sK3),sK1))) | (v1_xboole_0(k2_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) & ~v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK3,sK1)))) [skolemisation 22,33]
35. ~v1_xboole_0(k2_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | ~v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK3,sK1))) [cnf transformation 34] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK3,sK1))) | v1_xboole_0(k2_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) [cnf transformation 34] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
41. k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 25] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. ~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 27] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0) [cnf transformation 28] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0) [cnf transformation 29] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
46. v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0) [cnf transformation 30] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. ~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 32] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. 1 <=> v1_xboole_0(k2_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) [avatar definition]
50. ~v1_xboole_0(k2_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) <- (~1) [avatar component clause 49] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
51. v1_xboole_0(k2_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) <- (1) [avatar component clause 49] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
53. 2 <=> v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK3,sK1))) [avatar definition]
54. ~v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK3,sK1))) <- (~2) [avatar component clause 53] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
55. v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK3,sK1))) <- (2) [avatar component clause 53] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. 1 | 2 [avatar split clause 38,53,49]
57. ~2 | ~1 [avatar split clause 35,49,53]
59. ~v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) <- (~1) [resolution 42,50] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
61. 3 <=> v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) [avatar definition]
62. v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) <- (3) [avatar component clause 61] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
63. ~v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) <- (~3) [avatar component clause 61] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
64. ~3 | 1 [avatar split clause 59,49,61]
67. v1_xboole_0(k2_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 42,41] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
70. 4 <=> v1_xboole_0(k2_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) [avatar definition]
71. v1_xboole_0(k2_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) <- (4) [avatar component clause 70] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. ~v1_xboole_0(k2_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) <- (~4) [avatar component clause 70] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. 5 <=> v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK1,sK2),k2_gate_1(sK1,sK3))) [avatar definition]
77. ~v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK1,sK2),k2_gate_1(sK1,sK3))) <- (~5) [avatar component clause 76] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
78. v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK1,sK2),k2_gate_1(sK1,sK3))) <- (5) [avatar component clause 76] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
90. ~v1_xboole_0(sK1) <- (~4) [resolution 72,42] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
92. 6 <=> v1_xboole_0(sK1) [avatar definition]
94. ~v1_xboole_0(sK1) <- (~6) [avatar component clause 92] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. ~6 | 4 [avatar split clause 90,70,92]
96. ~v1_xboole_0(sK0) <- (~3) [resolution 44,63] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
98. 7 <=> v1_xboole_0(sK0) [avatar definition]
100. ~v1_xboole_0(sK0) <- (~7) [avatar component clause 98] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
101. ~7 | 3 [avatar split clause 96,61,98]
102. ~v1_xboole_0(sK3) <- (~3) [resolution 45,63] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
104. 8 <=> v1_xboole_0(sK3) [avatar definition]
106. ~v1_xboole_0(sK3) <- (~8) [avatar component clause 104] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
107. ~8 | 3 [avatar split clause 102,61,104]
108. ~v1_xboole_0(sK2) <- (~3) [resolution 46,63] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
110. 9 <=> v1_xboole_0(sK2) [avatar definition]
112. ~v1_xboole_0(sK2) <- (~9) [avatar component clause 110] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
113. ~9 | 3 [avatar split clause 108,61,110]
122. 10 <=> v1_xboole_0(k2_gate_1(sK0,sK1)) [avatar definition]
123. ~v1_xboole_0(k2_gate_1(sK0,sK1)) <- (~10) [avatar component clause 122] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. v1_xboole_0(k2_gate_1(sK0,sK1)) <- (10) [avatar component clause 122] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. 11 <=> v1_xboole_0(k2_gate_1(sK1,sK2)) [avatar definition]
127. ~v1_xboole_0(k2_gate_1(sK1,sK2)) <- (~11) [avatar component clause 126] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. v1_xboole_0(k2_gate_1(sK1,sK2)) <- (11) [avatar component clause 126] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
130. 12 <=> v1_xboole_0(k2_gate_1(sK1,sK3)) [avatar definition]
131. ~v1_xboole_0(k2_gate_1(sK1,sK3)) <- (~12) [avatar component clause 130] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
132. v1_xboole_0(k2_gate_1(sK1,sK3)) <- (12) [avatar component clause 130] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
135. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (10) [resolution 124,43] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
137. v1_xboole_0(sK0) <- (~6, 10) [subsumption resolution 135,94] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
140. 7 | 6 | ~10 [avatar split clause 137,122,92,98]
141. ~v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK1,sK3))) <- (~2) [forward demodulation 54,41] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
142. ~v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK1,sK2),k2_gate_1(sK1,sK3))) <- (~2) [forward demodulation 141,41] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
143. ~5 | 2 [avatar split clause 142,53,76]
144. v1_xboole_0(k2_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) <- (1) [forward demodulation 51,41] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
145. 4 | ~1 [avatar split clause 144,49,70]
146. v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) | v1_xboole_0(sK1) <- (4) [resolution 71,43] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
150. ~v1_xboole_0(k2_gate_1(sK1,sK2)) <- (~5) [resolution 77,46] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
151. ~v1_xboole_0(k2_gate_1(sK1,sK3)) <- (~5) [resolution 77,45] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
152. ~v1_xboole_0(k2_gate_1(sK0,sK1)) <- (~5) [resolution 77,44] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
156. ~10 | 5 [avatar split clause 152,76,122]
157. ~12 | 5 [avatar split clause 151,76,130]
158. ~11 | 5 [avatar split clause 150,76,126]
160. ~v1_xboole_0(sK0) <- (~10) [resolution 123,42] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
164. ~7 | 10 [avatar split clause 160,122,98]
165. ~v1_xboole_0(sK1) <- (~11) [resolution 127,42] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
168. v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK0) <- (3) [resolution 62,47] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
179. ~v1_xboole_0(sK2) <- (~11) [resolution 67,127] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
180. ~v1_xboole_0(sK3) <- (~12) [resolution 67,131] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
183. ~8 | 12 [avatar split clause 180,130,104]
184. ~9 | 11 [avatar split clause 179,126,110]
185. v1_xboole_0(sK3) | v1_xboole_0(sK2) <- (3, ~7) [subsumption resolution 168,100] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
186. 9 | 8 | ~3 | 7 [avatar split clause 185,98,61,104,110]
189. ~6 | 11 [avatar split clause 165,126,92]
190. 6 | 3 | ~4 [avatar split clause 146,70,61,92]
191. v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK2,sK1),k2_gate_1(sK1,sK3))) <- (2) [forward demodulation 55,41] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
192. v1_xboole_0(k8_gate_1(k2_gate_1(sK0,sK1),k2_gate_1(sK1,sK2),k2_gate_1(sK1,sK3))) <- (2) [forward demodulation 191,41] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
193. 5 | ~2 [avatar split clause 192,53,76]
194. ~v1_xboole_0(k2_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) <- (~1) [forward demodulation 50,41] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
195. ~4 | 1 [avatar split clause 194,49,70]
199. v1_xboole_0(k2_gate_1(sK1,sK3)) | v1_xboole_0(k2_gate_1(sK1,sK2)) | v1_xboole_0(k2_gate_1(sK0,sK1)) <- (5) [resolution 78,47] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
204. v1_xboole_0(sK3) | v1_xboole_0(sK1) <- (12) [resolution 132,43] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
206. v1_xboole_0(sK1) <- (~8, 12) [subsumption resolution 204,106] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
207. $false <- (~6, ~8, 12) [subsumption resolution 206,94] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
208. 6 | 8 | ~12 [avatar contradiction clause 207]
209. v1_xboole_0(k2_gate_1(sK1,sK3)) | v1_xboole_0(k2_gate_1(sK1,sK2)) <- (5, ~10) [subsumption resolution 199,123] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
210. 11 | 12 | ~5 | 10 [avatar split clause 209,122,76,130,126]
212. v1_xboole_0(sK2) | v1_xboole_0(sK1) <- (11) [resolution 128,43] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
214. v1_xboole_0(sK1) <- (~9, 11) [subsumption resolution 212,112] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
215. $false <- (~6, ~9, 11) [subsumption resolution 214,94] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
216. 6 | 9 | ~11 [avatar contradiction clause 215]
217. $false [avatar sat refutation 56,57,64,95,101,107,113,140,143,145,156,157,158,164,183,184,186,189,190,193,195,208,210,216] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:278,thDist:-278}
% SZS output end Proof for gate_2__t2_gate_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 138
% Active clauses: 35
% Passive clauses: 42
% Final active clauses: 18
% Final passive clauses: 4
% Main loop iterations started: 35
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 18
% Fw demodulations: 13
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 28
% Unit resulting resolution: 6
% Forward superposition: 4
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 13
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 49
% SAT solver binary clauses: 31
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

