i: [1,116,0,2,1,1,0,0] t1_trees_a
a: [3,180,0,2,1,116]
s: 180
i: [1,115,0,2,1,1,0,0] t1_trees_a
a: [3,185,0,2,1,115]
s: 185
i: [1,114,0,2,1,1,0,0] t1_trees_a
a: [3,190,0,2,1,114]
s: 190
i: [1,113,0,2,1,1,0,0] t1_trees_a
a: [3,195,0,2,1,113]
s: 195
i: [1,112,0,2,1,1,0,0] t1_trees_a
a: [3,200,0,2,1,112]
s: 200
i: [1,111,0,2,1,1,0,0] t1_trees_a
a: [3,205,0,2,1,111]
s: 205
i: [1,110,0,2,1,1,0,0] t1_trees_a
a: [3,210,0,2,1,110]
s: 210
i: [1,109,0,2,1,1,0,0] t1_trees_a
a: [3,215,0,2,1,109]
s: 215
i: [1,108,0,2,1,1,0,0] t1_trees_a
a: [3,220,0,2,1,108]
s: 220
i: [1,107,0,2,1,1,0,0] t1_trees_a
a: [3,225,0,2,1,107]
s: 225
i: [1,106,0,2,1,1,0,0] t1_trees_a
a: [3,230,0,2,1,106]
s: 230
i: [1,105,0,2,1,1,0,0] t1_trees_a
a: [3,235,0,2,1,105]
s: 235
i: [1,118,0,3,1,1,0,0] t1_trees_a
a: [3,170,0,3,1,118]
s: 170
i: [1,117,0,7,1,1,0,0] t1_trees_a
a: [3,175,0,7,1,117]
s: 175
i: [1,119,0,2,1,0,0,2] fc6_membered
a: [3,240,0,2,1,119]
s: 240
i: [1,120,0,2,1,0,0,255] fc3_membered
a: [3,245,0,2,1,120]
s: 245
i: [1,165,0,2,1,0,0,255] rc3_finseq_1
a: [3,260,0,2,1,165]
s: 260
i: [1,164,0,2,1,0,0,255] rc3_finseq_1
a: [3,265,0,2,1,164]
s: 265
i: [1,163,0,2,1,0,0,255] rc3_finseq_1
a: [3,270,0,2,1,163]
s: 270
i: [1,157,0,3,1,0,0,2] reflexivity_r1_tarski
s: 157
i: [1,158,0,3,1,0,0,1] reflexivity_r3_xboole_0
s: 158
i: [1,121,0,3,1,0,0,1] redefinition_k5_numbers
a: [3,250,0,3,1,121]
s: 250
r: 240
d: [2,272,0,2,1,2,23,240,250]
a: [3,276,0,2,1,272]
s: 276
i: [1,125,0,4,2,0,0,255] cc2_membered
s: 125
i: [1,126,0,4,2,0,0,255] cc3_membered
s: 126
i: [1,129,0,4,2,0,0,255] cc1_membered
s: 129
d: [2,278,1,2,1,1,40,129,276]
a: [3,282,1,2,1,278]
s: 282
i: [1,131,0,4,2,0,0,255] cc4_membered
s: 131
d: [2,290,1,2,1,1,40,131,245]
a: [3,294,1,2,1,290]
s: 294
i: [1,132,0,4,2,0,0,255] cc5_membered
s: 132
i: [1,156,0,4,1,0,0,3] existence_m1_subset_1
s: 156
d: [2,296,1,2,1,1,40,132,245]
a: [3,300,1,2,1,296]
s: 300
i: [1,122,0,4,1,0,0,255] dt_k5_numbers
a: [3,255,0,4,1,122]
s: 255
d: [2,284,2,2,1,1,40,282,125]
a: [3,288,2,2,1,284]
s: 288
i: [1,134,0,5,2,0,0,2] cc25_membered
s: 134
i: [1,135,0,6,2,0,0,255] cc8_membered
s: 135
d: [2,302,3,2,1,1,40,288,126]
a: [3,306,3,2,1,302]
s: 306
i: [1,136,0,6,2,0,0,2] cc11_membered
s: 136
d: [2,322,4,2,1,1,40,306,132]
a: [3,327,4,2,1,322]
s: 327
i: [1,159,0,6,2,0,0,1] symmetry_r3_xboole_0
s: 159
i: [1,154,0,7,2,0,0,2] t3_subset
s: 154
d: [2,323,4,2,1,1,40,306,131]
a: [3,332,4,2,1,323]
s: 332
i: [1,155,0,7,2,0,0,2] t3_subset
s: 155
d: [2,308,1,3,1,0,40,134,156]
a: [3,312,1,3,1,308]
s: 312
i: [1,123,0,8,3,0,0,255] cc19_membered
s: 123
i: [1,124,0,8,3,0,0,255] cc20_membered
s: 124
d: [2,342,1,3,1,1,40,154,255]
a: [3,346,1,3,1,342]
s: 346
i: [1,127,0,8,3,0,0,255] cc23_membered
s: 127
d: [2,351,2,3,1,1,40,312,129]
a: [3,355,2,3,1,351]
s: 355
i: [1,128,0,8,3,0,0,255] cc22_membered
s: 128
i: [1,130,0,8,3,0,0,3] cc24_membered
s: 130
d: [2,366,3,3,1,1,40,355,125]
a: [3,370,3,3,1,366]
s: 370
i: [1,133,0,8,3,0,0,255] cc21_membered
s: 133
d: [2,378,4,3,1,1,40,370,126]
a: [3,382,4,3,1,378]
s: 382
i: [1,144,0,8,4,0,0,255] cc6_finseq_1
s: 144
i: [1,160,0,16,7,0,0,1] dt_k7_finseq_1
s: 160
d: [2,399,1,6,3,1,40,144,225]
d: [2,403,1,4,2,2,21,399,230]
d: [2,404,1,2,1,3,21,403,235]
a: [3,408,1,2,1,404]
s: 408
i: [1,161,0,16,7,0,0,1] dt_k7_finseq_1
s: 161
d: [2,400,1,6,3,1,40,144,210]
d: [2,410,1,4,2,2,21,400,215]
d: [2,411,1,2,1,3,21,410,220]
a: [3,415,1,2,1,411]
s: 415
i: [1,162,0,16,7,0,0,1] dt_k7_finseq_1
s: 162
i: [1,148,0,19,8,0,0,1] t1_trees_1
s: 148
d: [2,401,1,6,3,1,40,144,195]
d: [2,417,1,4,2,2,21,401,200]
d: [2,418,1,2,1,3,21,417,205]
a: [3,422,1,2,1,418]
s: 422
i: [1,149,0,19,8,0,0,1] t1_trees_1
s: 149
d: [2,402,1,6,3,1,40,144,180]
d: [2,424,1,4,2,2,21,402,185]
d: [2,425,1,2,1,3,21,424,190]
a: [3,429,1,2,1,425]
s: 429
i: [1,150,0,19,8,0,0,1] t1_trees_1
s: 150
i: [1,151,0,22,8,0,0,1] t1_trees_1
s: 151
d: [2,431,1,16,7,1,44,160,175]
d: [2,432,1,14,6,2,21,431,195]
d: [2,433,1,12,5,3,21,432,200]
d: [2,434,1,10,4,4,21,433,205]
d: [2,435,1,8,3,5,21,434,180]
d: [2,436,1,6,2,6,21,435,185]
d: [2,437,1,4,1,7,21,436,190]
a: [3,441,1,4,1,437]
s: 441
i: [1,153,0,27,12,0,0,1] t1_trees_2
s: 153
d: [2,443,1,16,7,1,44,161,175]
d: [2,444,1,14,6,2,21,443,195]
d: [2,445,1,12,5,3,21,444,200]
d: [2,446,1,10,4,4,21,445,205]
d: [2,447,1,8,3,5,21,446,180]
d: [2,448,1,6,2,6,21,447,185]
d: [2,449,1,4,1,7,21,448,190]
a: [3,453,1,4,1,449]
s: 453
i: [1,166,0,29,10,0,0,1] t1_trees_1
s: 166
d: [2,510,0,25,9,0,21,166,162]
d: [2,511,0,21,8,0,21,510,161]
d: [2,512,0,17,7,0,21,511,160]
s: 512
d: [2,456,1,16,7,1,44,162,175]
d: [2,459,1,14,6,2,21,456,195]
d: [2,460,1,12,5,3,21,459,200]
d: [2,461,1,10,4,4,21,460,205]
d: [2,462,1,8,3,5,21,461,180]
d: [2,463,1,6,2,6,21,462,185]
d: [2,464,1,4,1,7,21,463,190]
a: [3,468,1,4,1,464]
s: 468
d: [2,514,1,17,7,1,44,512,175]
d: [2,521,1,15,6,2,21,514,195]
d: [2,522,1,13,5,3,21,521,200]
d: [2,523,1,11,4,4,21,522,205]
d: [2,524,1,9,3,5,21,523,180]
d: [2,525,1,7,2,6,21,524,185]
d: [2,526,1,5,1,7,21,525,190]
a: [3,530,1,5,1,526]
s: 530
d: [2,532,2,12,3,1,40,468,144]
d: [2,533,2,8,2,2,21,532,453]
d: [2,534,2,4,1,3,21,533,441]
a: [3,538,2,4,1,534]
s: 538
d: [2,314,1,4,1,0,40,135,156]
a: [3,319,1,4,1,314]
s: 319
d: [2,334,1,4,1,0,40,136,156]
a: [3,338,1,4,1,334]
s: 338
d: [2,480,2,24,11,1,40,153,346]
a: [3,485,2,2,1,480]
s: 485
d: [2,341,1,5,1,0,40,154,156]
s: 341
d: [2,387,5,3,1,1,40,382,132]
a: [3,392,5,3,1,387]
s: 392
d: [2,348,1,5,2,0,40,155,135]
s: 348
d: [2,349,1,5,2,0,40,155,136]
s: 349
d: [2,388,5,3,1,1,40,382,131]
a: [3,397,5,3,1,388]
s: 397
d: [2,357,1,6,2,0,40,123,156]
s: 357
d: [2,540,2,32,11,1,40,530,153]
d: [2,541,2,28,10,2,21,540,468]
d: [2,542,2,24,9,3,21,541,453]
d: [2,543,2,20,8,4,21,542,441]
d: [2,544,2,18,7,5,21,543,180]
d: [2,545,2,16,6,6,21,544,185]
d: [2,546,2,14,5,7,21,545,190]
s: 546
d: [2,571,3,21,10,7,40,546,512]
d: [2,575,3,15,7,7,18,571]
d: [2,576,3,12,6,8,21,575,170]
d: [2,577,3,10,5,9,21,576,225]
d: [2,578,3,8,4,10,21,577,230]
d: [2,579,3,6,3,11,21,578,235]
d: [2,580,3,4,2,12,21,579,210]
d: [2,581,3,2,1,13,21,580,215]
d: [2,582,3,0,0,14,21,581,220]
e: 582
f: 582,526,464,449,437,105,106,107,108,109,110,111,112,113,114,115,116,117,118
% Refutation found. Thanks to Tanya!
% SZS status Theorem for trees_a__t1_trees_a
% SZS output start Proof for trees_a__t1_trees_a
1. ! [X0] : ((v1_finseq_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) => (k7_finseq_1(X0,X1) = k7_finseq_1(X3,X2) => r3_xboole_0(X0,X3)))))) [input t1_trees_a]
2. ~! [X0] : ((v1_finseq_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) => (k7_finseq_1(X0,X1) = k7_finseq_1(X3,X2) => r3_xboole_0(X0,X3)))))) [negated conjecture 1]
31. ! [X0,X1] : ((v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1) & v1_finseq_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v1_finseq_1(k7_finseq_1(X0,X1)) & v1_funct_1(k7_finseq_1(X0,X1)) & v1_relat_1(k7_finseq_1(X0,X1)))) [input dt_k7_finseq_1]
43. ! [X0] : ((v1_finseq_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) <=> ? [X2] : (k7_finseq_1(X0,X2) = X1 & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2))))) [input t1_trees_1]
44. ! [X0] : ((v1_finseq_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_tarski(X1,X2) & r1_tarski(X0,X2)) => r3_xboole_0(X0,X1))))) [input t1_trees_2]
60. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r3_xboole_0(X0,X3) & k7_finseq_1(X0,X1) = k7_finseq_1(X3,X2)) & (v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3))) & (v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2))) & (v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1))) & (v1_finseq_1(X0) & v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
61. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r3_xboole_0(X0,X3) & k7_finseq_1(X0,X1) = k7_finseq_1(X3,X2) & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_finseq_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 60]
84. ! [X0] : (! [X1] : ((r1_tarski(X0,X1) <=> ? [X2] : (k7_finseq_1(X0,X2) = X1 & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2))) | (~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 43]
85. ! [X0] : (! [X1] : ((r1_tarski(X0,X1) <=> ? [X2] : (k7_finseq_1(X0,X2) = X1 & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2))) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 84]
86. ! [X0] : (! [X1] : (! [X2] : ((r3_xboole_0(X0,X1) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X2))) | (~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 44]
87. ! [X0] : (! [X1] : (! [X2] : (r3_xboole_0(X0,X1) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 86]
89. ! [X0,X1] : ((v1_finseq_1(k7_finseq_1(X0,X1)) & v1_funct_1(k7_finseq_1(X0,X1)) & v1_relat_1(k7_finseq_1(X0,X1))) | (~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 31]
90. ! [X0,X1] : ((v1_finseq_1(k7_finseq_1(X0,X1)) & v1_funct_1(k7_finseq_1(X0,X1)) & v1_relat_1(k7_finseq_1(X0,X1))) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 89]
91. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r3_xboole_0(X0,X3) & k7_finseq_1(X0,X1) = k7_finseq_1(X3,X2) & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_finseq_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r3_xboole_0(sK0,X3) & k7_finseq_1(X3,X2) = k7_finseq_1(sK0,X1) & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_finseq_1(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
92. ? [X1] : (? [X2] : (? [X3] : (~r3_xboole_0(sK0,X3) & k7_finseq_1(X3,X2) = k7_finseq_1(sK0,X1) & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (~r3_xboole_0(sK0,X3) & k7_finseq_1(X3,X2) = k7_finseq_1(sK0,sK1) & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_finseq_1(sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
93. ? [X2] : (? [X3] : (~r3_xboole_0(sK0,X3) & k7_finseq_1(X3,X2) = k7_finseq_1(sK0,sK1) & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => (? [X3] : (~r3_xboole_0(sK0,X3) & k7_finseq_1(sK0,sK1) = k7_finseq_1(X3,sK2) & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_finseq_1(sK2) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
94. ? [X3] : (~r3_xboole_0(sK0,X3) & k7_finseq_1(sK0,sK1) = k7_finseq_1(X3,sK2) & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) => (~r3_xboole_0(sK0,sK3) & k7_finseq_1(sK0,sK1) = k7_finseq_1(sK3,sK2) & v1_finseq_1(sK3) & v1_funct_1(sK3) & v1_relat_1(sK3)) [choice axiom]
95. (((~r3_xboole_0(sK0,sK3) & k7_finseq_1(sK0,sK1) = k7_finseq_1(sK3,sK2) & v1_finseq_1(sK3) & v1_funct_1(sK3) & v1_relat_1(sK3)) & v1_finseq_1(sK2) & v1_funct_1(sK2) & v1_relat_1(sK2)) & v1_finseq_1(sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_finseq_1(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 61,94,93,92,91]
96. ! [X0] : (! [X1] : (((r1_tarski(X0,X1) | ! [X2] : (k7_finseq_1(X0,X2) != X1 | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X2] : (k7_finseq_1(X0,X2) = X1 & v1_finseq_1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) | ~r1_tarski(X0,X1))) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 85]
97. ! [X0] : (! [X1] : (((r1_tarski(X0,X1) | ! [X2] : (k7_finseq_1(X0,X2) != X1 | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X3] : (k7_finseq_1(X0,X3) = X1 & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) | ~r1_tarski(X0,X1))) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [rectify 96]
98. ! [X1,X0] : (? [X3] : (k7_finseq_1(X0,X3) = X1 & v1_finseq_1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) => (k7_finseq_1(X0,sK4(X0,X1)) = X1 & v1_finseq_1(sK4(X0,X1)) & v1_funct_1(sK4(X0,X1)) & v1_relat_1(sK4(X0,X1)))) [choice axiom]
99. ! [X0] : (! [X1] : (((r1_tarski(X0,X1) | ! [X2] : (k7_finseq_1(X0,X2) != X1 | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & ((k7_finseq_1(X0,sK4(X0,X1)) = X1 & v1_finseq_1(sK4(X0,X1)) & v1_funct_1(sK4(X0,X1)) & v1_relat_1(sK4(X0,X1))) | ~r1_tarski(X0,X1))) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [skolemisation 97,98]
105. v1_relat_1(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
106. v1_funct_1(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
107. v1_finseq_1(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
108. v1_relat_1(sK1) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
109. v1_funct_1(sK1) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
110. v1_finseq_1(sK1) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
111. v1_relat_1(sK2) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
112. v1_funct_1(sK2) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
113. v1_finseq_1(sK2) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
114. v1_relat_1(sK3) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
115. v1_funct_1(sK3) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
116. v1_finseq_1(sK3) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
117. k7_finseq_1(sK0,sK1) = k7_finseq_1(sK3,sK2) [cnf transformation 95] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
118. ~r3_xboole_0(sK0,sK3) [cnf transformation 95] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
152. r1_tarski(X0,X1) | k7_finseq_1(X0,X2) != X1 | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 99] {a:0,w:26,wCS:10400,sine:1,thAx:0,allAx:2,thDist:-2}
153. ~r1_tarski(X1,X2) | r3_xboole_0(X0,X1) | ~r1_tarski(X0,X2) | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 87] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
160. v1_relat_1(k7_finseq_1(X0,X1)) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 90] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
161. v1_funct_1(k7_finseq_1(X0,X1)) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 90] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
162. v1_finseq_1(k7_finseq_1(X0,X1)) | ~v1_finseq_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 90] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
166. r1_tarski(X0,k7_finseq_1(X0,X2)) | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_finseq_1(k7_finseq_1(X0,X2)) | ~v1_funct_1(k7_finseq_1(X0,X2)) | ~v1_relat_1(k7_finseq_1(X0,X2)) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [equality resolution 152] {a:0,w:29,wCS:11600,sine:1,thAx:0,allAx:2,thDist:-2}
168. 1 <=> r3_xboole_0(sK0,sK3) [avatar definition]
170. ~r3_xboole_0(sK0,sK3) <- (~1) [avatar component clause 168] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
171. ~1 [avatar split clause 118,168]
173. 2 <=> k7_finseq_1(sK0,sK1) = k7_finseq_1(sK3,sK2) [avatar definition]
175. k7_finseq_1(sK0,sK1) = k7_finseq_1(sK3,sK2) <- (2) [avatar component clause 173] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
176. 2 [avatar split clause 117,173]
178. 3 <=> v1_finseq_1(sK3) [avatar definition]
180. v1_finseq_1(sK3) <- (3) [avatar component clause 178] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
181. 3 [avatar split clause 116,178]
183. 4 <=> v1_funct_1(sK3) [avatar definition]
185. v1_funct_1(sK3) <- (4) [avatar component clause 183] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
186. 4 [avatar split clause 115,183]
188. 5 <=> v1_relat_1(sK3) [avatar definition]
190. v1_relat_1(sK3) <- (5) [avatar component clause 188] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
191. 5 [avatar split clause 114,188]
193. 6 <=> v1_finseq_1(sK2) [avatar definition]
195. v1_finseq_1(sK2) <- (6) [avatar component clause 193] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
196. 6 [avatar split clause 113,193]
198. 7 <=> v1_funct_1(sK2) [avatar definition]
200. v1_funct_1(sK2) <- (7) [avatar component clause 198] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
201. 7 [avatar split clause 112,198]
203. 8 <=> v1_relat_1(sK2) [avatar definition]
205. v1_relat_1(sK2) <- (8) [avatar component clause 203] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
206. 8 [avatar split clause 111,203]
208. 9 <=> v1_finseq_1(sK1) [avatar definition]
210. v1_finseq_1(sK1) <- (9) [avatar component clause 208] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
211. 9 [avatar split clause 110,208]
213. 10 <=> v1_funct_1(sK1) [avatar definition]
215. v1_funct_1(sK1) <- (10) [avatar component clause 213] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
216. 10 [avatar split clause 109,213]
218. 11 <=> v1_relat_1(sK1) [avatar definition]
220. v1_relat_1(sK1) <- (11) [avatar component clause 218] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
221. 11 [avatar split clause 108,218]
223. 12 <=> v1_finseq_1(sK0) [avatar definition]
225. v1_finseq_1(sK0) <- (12) [avatar component clause 223] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
226. 12 [avatar split clause 107,223]
228. 13 <=> v1_funct_1(sK0) [avatar definition]
230. v1_funct_1(sK0) <- (13) [avatar component clause 228] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
231. 13 [avatar split clause 106,228]
233. 14 <=> v1_relat_1(sK0) [avatar definition]
235. v1_relat_1(sK0) <- (14) [avatar component clause 233] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
236. 14 [avatar split clause 105,233]
431. v1_relat_1(k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK2) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2) [superposition 160,175] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
432. v1_relat_1(k7_finseq_1(sK0,sK1)) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6) [subsumption resolution 431,195] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
433. v1_relat_1(k7_finseq_1(sK0,sK1)) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6, 7) [subsumption resolution 432,200] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
434. v1_relat_1(k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6, 7, 8) [subsumption resolution 433,205] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
435. v1_relat_1(k7_finseq_1(sK0,sK1)) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 3, 6, 7, 8) [subsumption resolution 434,180] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
436. v1_relat_1(k7_finseq_1(sK0,sK1)) | ~v1_relat_1(sK3) <- (2, 3, 4, 6, 7, 8) [subsumption resolution 435,185] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
437. v1_relat_1(k7_finseq_1(sK0,sK1)) <- (2, 3, 4, 5, 6, 7, 8) [subsumption resolution 436,190] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
439. 43 <=> v1_relat_1(k7_finseq_1(sK0,sK1)) [avatar definition]
441. v1_relat_1(k7_finseq_1(sK0,sK1)) <- (43) [avatar component clause 439] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
442. 43 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 437,203,198,193,188,183,178,173,439]
443. v1_funct_1(k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK2) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2) [superposition 161,175] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
444. v1_funct_1(k7_finseq_1(sK0,sK1)) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6) [subsumption resolution 443,195] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
445. v1_funct_1(k7_finseq_1(sK0,sK1)) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6, 7) [subsumption resolution 444,200] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
446. v1_funct_1(k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6, 7, 8) [subsumption resolution 445,205] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
447. v1_funct_1(k7_finseq_1(sK0,sK1)) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 3, 6, 7, 8) [subsumption resolution 446,180] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
448. v1_funct_1(k7_finseq_1(sK0,sK1)) | ~v1_relat_1(sK3) <- (2, 3, 4, 6, 7, 8) [subsumption resolution 447,185] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
449. v1_funct_1(k7_finseq_1(sK0,sK1)) <- (2, 3, 4, 5, 6, 7, 8) [subsumption resolution 448,190] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
451. 44 <=> v1_funct_1(k7_finseq_1(sK0,sK1)) [avatar definition]
453. v1_funct_1(k7_finseq_1(sK0,sK1)) <- (44) [avatar component clause 451] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
454. 44 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 449,203,198,193,188,183,178,173,451]
456. v1_finseq_1(k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK2) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2) [superposition 162,175] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
459. v1_finseq_1(k7_finseq_1(sK0,sK1)) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6) [subsumption resolution 456,195] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
460. v1_finseq_1(k7_finseq_1(sK0,sK1)) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6, 7) [subsumption resolution 459,200] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
461. v1_finseq_1(k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6, 7, 8) [subsumption resolution 460,205] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
462. v1_finseq_1(k7_finseq_1(sK0,sK1)) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 3, 6, 7, 8) [subsumption resolution 461,180] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
463. v1_finseq_1(k7_finseq_1(sK0,sK1)) | ~v1_relat_1(sK3) <- (2, 3, 4, 6, 7, 8) [subsumption resolution 462,185] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
464. v1_finseq_1(k7_finseq_1(sK0,sK1)) <- (2, 3, 4, 5, 6, 7, 8) [subsumption resolution 463,190] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
466. 45 <=> v1_finseq_1(k7_finseq_1(sK0,sK1)) [avatar definition]
468. v1_finseq_1(k7_finseq_1(sK0,sK1)) <- (45) [avatar component clause 466] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
469. 45 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 464,203,198,193,188,183,178,173,466]
510. r1_tarski(X0,k7_finseq_1(X0,X2)) | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(k7_finseq_1(X0,X2)) | ~v1_relat_1(k7_finseq_1(X0,X2)) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 166,162] {a:0,w:25,wCS:10000,sine:1,thAx:0,allAx:2,thDist:-2}
511. r1_tarski(X0,k7_finseq_1(X0,X2)) | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(k7_finseq_1(X0,X2)) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 510,161] {a:0,w:21,wCS:8400,sine:1,thAx:0,allAx:2,thDist:-2}
512. r1_tarski(X0,k7_finseq_1(X0,X2)) | ~v1_finseq_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 511,160] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
514. r1_tarski(sK3,k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK2) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2) [superposition 512,175] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
521. r1_tarski(sK3,k7_finseq_1(sK0,sK1)) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6) [subsumption resolution 514,195] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
522. r1_tarski(sK3,k7_finseq_1(sK0,sK1)) | ~v1_relat_1(sK2) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6, 7) [subsumption resolution 521,200] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
523. r1_tarski(sK3,k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 6, 7, 8) [subsumption resolution 522,205] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
524. r1_tarski(sK3,k7_finseq_1(sK0,sK1)) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, 3, 6, 7, 8) [subsumption resolution 523,180] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
525. r1_tarski(sK3,k7_finseq_1(sK0,sK1)) | ~v1_relat_1(sK3) <- (2, 3, 4, 6, 7, 8) [subsumption resolution 524,185] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
526. r1_tarski(sK3,k7_finseq_1(sK0,sK1)) <- (2, 3, 4, 5, 6, 7, 8) [subsumption resolution 525,190] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
528. 53 <=> r1_tarski(sK3,k7_finseq_1(sK0,sK1)) [avatar definition]
530. r1_tarski(sK3,k7_finseq_1(sK0,sK1)) <- (53) [avatar component clause 528] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
531. 53 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 526,203,198,193,188,183,178,173,528]
540. r3_xboole_0(X0,sK3) | ~r1_tarski(X0,k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(k7_finseq_1(sK0,sK1)) | ~v1_funct_1(k7_finseq_1(sK0,sK1)) | ~v1_relat_1(k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (53) [resolution 530,153] {a:2,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
541. r3_xboole_0(X0,sK3) | ~r1_tarski(X0,k7_finseq_1(sK0,sK1)) | ~v1_funct_1(k7_finseq_1(sK0,sK1)) | ~v1_relat_1(k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (45, 53) [subsumption resolution 540,468] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
542. r3_xboole_0(X0,sK3) | ~r1_tarski(X0,k7_finseq_1(sK0,sK1)) | ~v1_relat_1(k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (44, 45, 53) [subsumption resolution 541,453] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
543. r3_xboole_0(X0,sK3) | ~r1_tarski(X0,k7_finseq_1(sK0,sK1)) | ~v1_finseq_1(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (43, 44, 45, 53) [subsumption resolution 542,441] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
544. r3_xboole_0(X0,sK3) | ~r1_tarski(X0,k7_finseq_1(sK0,sK1)) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (3, 43, 44, 45, 53) [subsumption resolution 543,180] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
545. r3_xboole_0(X0,sK3) | ~r1_tarski(X0,k7_finseq_1(sK0,sK1)) | ~v1_relat_1(sK3) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (3, 4, 43, 44, 45, 53) [subsumption resolution 544,185] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
546. ~r1_tarski(X0,k7_finseq_1(sK0,sK1)) | r3_xboole_0(X0,sK3) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (3, 4, 5, 43, 44, 45, 53) [subsumption resolution 545,190] {a:2,w:14,wCS:1400,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
571. r3_xboole_0(sK0,sK3) | ~v1_finseq_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_finseq_1(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_finseq_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3, 4, 5, 43, 44, 45, 53) [resolution 546,512] {a:3,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
575. r3_xboole_0(sK0,sK3) | ~v1_finseq_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_finseq_1(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (3, 4, 5, 43, 44, 45, 53) [duplicate literal removal 571] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
576. ~v1_finseq_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_finseq_1(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 43, 44, 45, 53) [subsumption resolution 575,170] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
577. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_finseq_1(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 12, 43, 44, 45, 53) [subsumption resolution 576,225] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
578. ~v1_relat_1(sK0) | ~v1_finseq_1(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 12, 13, 43, 44, 45, 53) [subsumption resolution 577,230] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
579. ~v1_finseq_1(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 12, 13, 14, 43, 44, 45, 53) [subsumption resolution 578,235] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
580. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 9, 12, 13, 14, 43, 44, 45, 53) [subsumption resolution 579,210] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
581. ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 9, 10, 12, 13, 14, 43, 44, 45, 53) [subsumption resolution 580,215] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
582. $false <- (~1, 3, 4, 5, 9, 10, 11, 12, 13, 14, 43, 44, 45, 53) [subsumption resolution 581,220] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
583. 1 | ~3 | ~4 | ~5 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | ~43 | ~44 | ~45 | ~53 [avatar contradiction clause 582]
587. $false [avatar sat refutation 171,176,181,186,191,196,201,206,211,216,221,226,231,236,442,454,469,531,583] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:348,thDist:-348}
% SZS output end Proof for trees_a__t1_trees_a
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 44
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 9
% 
% >>> Saturation
% Initial clauses: 61
% Generated clauses: 252
% Active clauses: 83
% Passive clauses: 103
% Final active clauses: 82
% Final passive clauses: 19
% Main loop iterations started: 84
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 60
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 64
% Forward superposition: 5
% Backward superposition: 3
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 7
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 55
% SAT solver unit clauses: 24
% SAT solver binary clauses: 17
% 
% Memory used [KB]: 9722
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s ( own 0.0000 s ) 
% sine selection: 0.001 s
% forward subsumption: 0.001 s
% backward subsumption: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

