i: [1,1236,0,53,10,1,0,0] t42_gate_1
a: [3,1318,0,2,1,1236]
s: 1318
a: [3,1327,0,2,1,1236]
s: 1327
a: [3,1335,0,2,1,1236]
s: 1335
a: [3,1343,0,2,1,1236]
s: 1343
a: [3,1350,0,2,1,1236]
s: 1350
i: [1,210,0,9,2,1,0,0] t42_gate_1
a: [3,1397,0,2,1,210]
s: 1397
i: [1,208,0,8,2,1,0,0] t42_gate_1
a: [3,1406,0,2,1,208]
s: 1406
i: [1,206,0,7,2,1,0,0] t42_gate_1
a: [3,1417,0,2,1,206]
s: 1417
i: [1,204,0,6,2,1,0,0] t42_gate_1
a: [3,1426,0,2,1,204]
s: 1426
i: [1,200,0,8,2,1,0,0] t42_gate_1
a: [3,1443,0,2,1,200]
s: 1443
i: [1,198,0,8,2,1,0,0] t42_gate_1
a: [3,1452,0,2,1,198]
s: 1452
i: [1,196,0,7,2,1,0,0] t42_gate_1
a: [3,1462,0,2,1,196]
s: 1462
i: [1,194,0,6,2,1,0,0] t42_gate_1
a: [3,1473,0,2,1,194]
s: 1473
i: [1,190,0,7,2,1,0,0] t42_gate_1
a: [3,1488,0,2,1,190]
s: 1488
i: [1,188,0,7,2,1,0,0] t42_gate_1
a: [3,1498,0,2,1,188]
s: 1498
i: [1,186,0,6,2,1,0,0] t42_gate_1
a: [3,1509,0,2,1,186]
s: 1509
i: [1,182,0,6,2,1,0,0] t42_gate_1
a: [3,1525,0,2,1,182]
s: 1525
i: [1,180,0,6,2,1,0,0] t42_gate_1
a: [3,1534,0,2,1,180]
s: 1534
i: [1,176,0,5,2,1,0,0] t42_gate_1
a: [3,1550,0,2,1,176]
s: 1550
i: [1,174,0,5,2,1,0,0] t42_gate_1
a: [3,1561,0,2,1,174]
s: 1561
i: [1,172,0,6,2,1,0,0] t42_gate_1
a: [3,1571,0,2,1,172]
s: 1571
i: [1,170,0,6,2,1,0,0] t42_gate_1
a: [3,1580,0,2,1,170]
s: 1580
i: [1,168,0,6,2,1,0,0] t42_gate_1
a: [3,1591,0,2,1,168]
s: 1591
i: [1,166,0,6,2,1,0,0] t42_gate_1
a: [3,1601,0,2,1,166]
s: 1601
i: [1,164,0,6,2,1,0,0] t42_gate_1
a: [3,1611,0,2,1,164]
s: 1611
i: [1,162,0,6,2,1,0,0] t42_gate_1
a: [3,1620,0,2,1,162]
s: 1620
i: [1,160,0,6,2,1,0,0] t42_gate_1
a: [3,1631,0,2,1,160]
s: 1631
i: [1,158,0,6,2,1,0,0] t42_gate_1
a: [3,1640,0,2,1,158]
s: 1640
i: [1,156,0,6,2,1,0,0] t42_gate_1
a: [3,1651,0,2,1,156]
s: 1651
i: [1,154,0,6,2,1,0,0] t42_gate_1
a: [3,1661,0,2,1,154]
s: 1661
i: [1,152,0,6,2,1,0,0] t42_gate_1
a: [3,1671,0,2,1,152]
s: 1671
i: [1,150,0,6,2,1,0,0] t42_gate_1
a: [3,1681,0,2,1,150]
s: 1681
a: [3,1547,0,3,1,176]
s: 1547
a: [3,1556,0,3,1,174]
s: 1556
a: [3,1423,0,4,1,204]
s: 1423
i: [1,202,0,6,2,1,0,0] t42_gate_1
a: [3,1433,0,4,1,202]
s: 1433
a: [3,1468,0,4,1,194]
s: 1468
i: [1,192,0,6,2,1,0,0] t42_gate_1
a: [3,1479,0,4,1,192]
s: 1479
a: [3,1504,0,4,1,186]
s: 1504
i: [1,184,0,6,2,1,0,0] t42_gate_1
a: [3,1515,0,4,1,184]
s: 1515
a: [3,1520,0,4,1,182]
s: 1520
a: [3,1531,0,4,1,180]
s: 1531
i: [1,178,0,6,2,1,0,0] t42_gate_1
a: [3,1540,0,4,1,178]
s: 1540
a: [3,1566,0,4,1,172]
s: 1566
a: [3,1577,0,4,1,170]
s: 1577
a: [3,1586,0,4,1,168]
s: 1586
a: [3,1596,0,4,1,166]
s: 1596
a: [3,1606,0,4,1,164]
s: 1606
a: [3,1617,0,4,1,162]
s: 1617
a: [3,1626,0,4,1,160]
s: 1626
a: [3,1637,0,4,1,158]
s: 1637
a: [3,1646,0,4,1,156]
s: 1646
a: [3,1656,0,4,1,154]
s: 1656
a: [3,1666,0,4,1,152]
s: 1666
a: [3,1676,0,4,1,150]
s: 1676
a: [3,1315,0,5,1,1236]
s: 1315
a: [3,1412,0,5,1,206]
s: 1412
a: [3,1459,0,5,1,196]
s: 1459
a: [3,1485,0,5,1,190]
s: 1485
a: [3,1495,0,5,1,188]
s: 1495
a: [3,1403,0,6,1,208]
s: 1403
a: [3,1438,0,6,1,200]
s: 1438
a: [3,1449,0,6,1,198]
s: 1449
a: [3,1323,0,7,1,1236]
s: 1323
i: [1,212,0,9,2,1,0,0] t42_gate_1
a: [3,1386,0,7,1,212]
s: 1386
a: [3,1392,0,7,1,210]
s: 1392
i: [1,1237,0,2,1,0,0,2] fc1_xboole_0
a: [3,1687,0,2,1,1237]
s: 1687
i: [1,1310,0,2,1,0,0,1] rc2_xboole_0
a: [3,1692,0,2,1,1310]
s: 1692
i: [1,1311,0,2,1,0,0,1] rc1_xboole_0
a: [3,1697,0,2,1,1311]
s: 1697
i: [1,1238,0,5,2,0,0,1] fc2_gate_1
s: 1238
r: 1556
r: 1561
a: [3,1560,0,2,1,174]
s: 1560
a: [3,1557,0,3,1,174]
s: 1557
i: [1,1239,0,5,2,0,0,1] fc1_gate_1
s: 1239
d: [2,1705,1,0,0,2,52,1560,1547,1239]
e: 1705
d: [2,1706,1,2,1,1,40,1239,1547]
d: [2,1710,1,0,0,2,21,1706,1560]
e: 1710
r: 1547
r: 1550
a: [3,1551,0,2,1,176]
s: 1551
a: [3,1546,0,3,1,176]
s: 1546
i: [1,1283,0,6,2,0,0,1] fc16_gate_1
s: 1283
r: 1586
r: 1591
r: 1646
r: 1651
r: 1676
r: 1681
a: [3,1590,0,2,1,168]
s: 1590
a: [3,1650,0,2,1,156]
s: 1650
a: [3,1680,0,2,1,150]
s: 1680
a: [3,1587,0,4,1,168]
s: 1587
a: [3,1647,0,4,1,156]
s: 1647
d: [2,1699,1,2,1,1,40,1238,1556]
a: [3,1702,1,2,1,1699]
s: 1702
a: [3,1677,0,4,1,150]
s: 1677
i: [1,1284,0,6,2,0,0,1] fc13_gate_1
s: 1284
i: [1,1285,0,6,2,0,0,1] fc4_gate_1
s: 1285
d: [2,1737,1,0,0,2,52,1650,1468,1285]
e: 1737
d: [2,1738,1,0,0,2,52,1680,1504,1285]
e: 1738
d: [2,1739,1,2,1,1,40,1285,1468]
d: [2,1743,1,0,0,2,21,1739,1650]
e: 1743
d: [2,1740,1,2,1,1,40,1285,1504]
d: [2,1745,1,0,0,2,21,1740,1680]
e: 1745
r: 1468
r: 1473
r: 1504
r: 1509
a: [3,1472,0,2,1,194]
s: 1472
a: [3,1508,0,2,1,186]
s: 1508
a: [3,1469,0,4,1,194]
s: 1469
d: [2,1715,1,2,1,1,40,1283,1646]
a: [3,1725,1,2,1,1715]
s: 1725
a: [3,1505,0,4,1,186]
s: 1505
i: [1,1278,0,7,1,0,0,1] commutativity_k34_gate_1
s: 1278
d: [2,1716,1,2,1,1,40,1283,1586]
a: [3,1730,1,2,1,1716]
s: 1730
i: [1,1279,0,7,1,0,0,1] commutativity_k6_gate_1
s: 1279
d: [2,1747,1,6,2,0,44,1284,1279]
s: 1747
i: [1,1280,0,7,1,0,0,1] commutativity_k2_gate_1
s: 1280
r: 1531
r: 1534
r: 1520
r: 1525
a: [3,1524,0,2,1,182]
s: 1524
a: [3,1535,0,2,1,180]
s: 1535
a: [3,1521,0,4,1,182]
s: 1521
d: [2,1749,2,2,1,1,40,1747,1577]
a: [3,1756,2,2,1,1749]
s: 1756
d: [2,1763,0,4,1,1,23,1531,1280]
a: [3,1768,0,4,1,1763]
s: 1768
i: [1,1281,0,7,1,0,0,1] commutativity_k4_gate_1
s: 1281
r: 1540
r: 1515
r: 1479
r: 1433
r: 1327
r: 1335
r: 1343
a: [3,1326,0,2,1,1236]
s: 1326
a: [3,1334,0,2,1,1236]
s: 1334
a: [3,1342,0,2,1,1236]
s: 1342
d: [2,1775,0,4,1,1,23,1540,1281]
a: [3,1782,0,4,1,1775]
s: 1782
d: [2,1776,0,4,1,1,23,1515,1281]
a: [3,1786,0,4,1,1776]
s: 1786
d: [2,1750,2,2,1,1,40,1747,1637]
a: [3,1761,2,2,1,1750]
s: 1761
d: [2,1777,0,4,1,1,23,1479,1281]
a: [3,1791,0,4,1,1777]
s: 1791
d: [2,1778,0,4,1,1,23,1433,1281]
a: [3,1796,0,4,1,1778]
s: 1796
i: [1,1282,0,7,1,0,0,1] commutativity_k7_gate_1
s: 1282
r: 1521
r: 1524
r: 1535
r: 1768
s: 1525
s: 1534
a: [3,1769,0,4,1,1763]
s: 1769
d: [2,1805,0,4,1,1,23,1521,1282]
a: [3,1810,0,4,1,1805]
s: 1810
i: [1,1292,0,7,2,0,0,1] fc19_gate_1
s: 1292
d: [2,1818,1,0,0,2,52,1580,1412,1292]
e: 1818
d: [2,1819,1,2,1,1,40,1292,1412]
d: [2,1821,1,0,0,2,21,1819,1580]
e: 1821
r: 1577
r: 1580
a: [3,1581,0,2,1,170]
s: 1581
a: [3,1576,0,4,1,170]
s: 1576
i: [1,1293,0,7,2,0,0,1] fc20_gate_1
s: 1293
d: [2,1748,1,6,2,0,44,1284,1279]
s: 1748
i: [1,1294,0,7,2,0,0,1] fc18_gate_1
s: 1294
d: [2,1825,1,0,0,2,52,1650,1412,1294]
e: 1825
d: [2,1826,1,2,1,1,40,1294,1412]
d: [2,1828,1,0,0,2,21,1826,1650]
e: 1828
r: 1637
r: 1640
r: 1647
r: 1650
r: 1725
a: [3,1641,0,2,1,158]
s: 1641
s: 1651
a: [3,1636,0,4,1,158]
s: 1636
a: [3,1726,1,2,1,1715]
s: 1726
s: 1646
i: [1,1247,0,8,2,0,0,1] d13_gate_1
s: 1247
d: [2,1764,1,6,2,0,44,1285,1280]
s: 1764
i: [1,1248,0,8,2,0,0,1] d13_gate_1
s: 1248
d: [2,1765,1,6,2,0,44,1285,1280]
s: 1765
i: [1,1249,0,8,2,0,0,1] d13_gate_1
s: 1249
i: [1,1286,0,8,3,0,0,1] fc3_gate_1
s: 1286
d: [2,1837,1,0,0,3,52,1651,1611,1469,1286]
e: 1837
d: [2,1842,1,4,2,1,40,1286,1469]
d: [2,1847,1,2,1,2,21,1842,1611]
d: [2,1848,1,0,0,3,21,1847,1651]
e: 1848
r: 1469
r: 1472
s: 1473
s: 1468
d: [2,1806,1,6,2,0,44,1283,1282]
s: 1806
d: [2,1852,3,0,0,2,52,1761,1646,1806]
e: 1852
d: [2,1854,2,2,1,1,40,1806,1646]
d: [2,1858,2,0,0,2,21,1854,1761]
e: 1858
r: 1761
i: [1,1287,0,8,3,0,0,1] fc14_gate_1
s: 1287
d: [2,1861,3,0,0,3,52,1730,1756,1576,1287]
e: 1861
d: [2,1867,1,4,2,1,40,1287,1576]
d: [2,1871,1,2,1,2,21,1867,1756]
d: [2,1872,1,0,0,3,21,1871,1730]
e: 1872
r: 1730
r: 1606
r: 1611
r: 1412
r: 1417
r: 1468
r: 1473
a: [3,1416,0,2,1,206]
s: 1416
s: 1472
a: [3,1610,0,2,1,164]
s: 1610
d: [2,1714,1,2,1,1,40,1283,1676]
a: [3,1721,1,2,1,1714]
s: 1721
s: 1469
a: [3,1607,0,4,1,164]
s: 1607
a: [3,1731,1,2,1,1716]
s: 1731
a: [3,1413,0,5,1,206]
s: 1413
d: [2,1866,1,4,2,1,40,1287,1606]
a: [3,1877,1,2,1,1866]
s: 1877
i: [1,1288,0,8,3,0,0,1] fc9_gate_1
s: 1288
d: [2,1896,1,0,0,3,52,1661,1551,1782,1288]
e: 1896
d: [2,1897,1,0,0,3,52,1631,1525,1786,1288]
e: 1897
d: [2,1898,1,0,0,3,52,1571,1443,1796,1288]
e: 1898
d: [2,1899,1,4,2,1,40,1288,1782]
d: [2,1908,1,2,1,2,21,1899,1551]
d: [2,1909,1,0,0,3,21,1908,1661]
e: 1909
d: [2,1900,1,4,2,1,40,1288,1786]
d: [2,1911,1,2,1,2,21,1900,1525]
d: [2,1912,1,0,0,3,21,1911,1631]
e: 1912
d: [2,1902,1,4,2,1,40,1288,1796]
d: [2,1914,1,2,1,2,21,1902,1443]
d: [2,1915,1,0,0,3,21,1914,1571]
e: 1915
r: 1318
r: 1326
r: 1342
r: 1782
r: 1786
r: 1796
a: [3,1319,0,2,1,1236]
s: 1319
s: 1327
s: 1343
a: [3,1881,1,2,1,1866]
s: 1881
a: [3,1781,0,4,1,1775]
s: 1781
a: [3,1787,0,4,1,1776]
s: 1787
a: [3,1760,2,2,1,1750]
s: 1760
a: [3,1797,0,4,1,1778]
s: 1797
a: [3,1331,0,9,1,1236]
s: 1331
i: [1,1289,0,8,3,0,0,1] fc7_gate_1
s: 1289
i: [1,1290,0,8,3,0,0,1] fc15_gate_1
s: 1290
d: [2,1864,1,4,2,1,40,1287,1666]
a: [3,1886,1,2,1,1864]
s: 1886
i: [1,1291,0,8,3,0,0,1] fc8_gate_1
s: 1291
d: [2,1807,1,6,2,0,44,1283,1282]
s: 1807
i: [1,1244,0,9,3,0,0,1] d34_gate_1
s: 1244
d: [2,1949,1,6,3,1,44,1626,1244]
d: [2,1952,1,4,2,2,21,1949,1687]
d: [2,1953,1,2,1,3,21,1952,1760]
d: [2,1954,1,0,0,4,21,1953,1726]
e: 1954
r: 1646
r: 1651
r: 1726
s: 1650
s: 1725
s: 1647
d: [2,1927,1,8,3,0,44,1289,1281]
s: 1927
i: [1,1245,0,9,3,0,0,1] d34_gate_1
s: 1245
d: [2,1973,1,6,3,1,44,1596,1245]
d: [2,1975,1,4,2,2,21,1973,1687]
d: [2,1976,1,2,1,3,21,1975,1881]
d: [2,1977,1,0,0,4,21,1976,1877]
e: 1977
r: 1596
r: 1601
r: 1334
r: 1791
s: 1335
a: [3,1600,0,2,1,166]
s: 1600
a: [3,1597,0,4,1,166]
s: 1597
a: [3,1792,0,4,1,1777]
s: 1792
a: [3,1347,0,11,1,1236]
s: 1347
i: [1,1251,0,9,2,0,0,1] d14_gate_1
s: 1251
d: [2,1928,1,8,3,0,44,1289,1281]
s: 1928
i: [1,1252,0,9,2,0,0,1] d14_gate_1
s: 1252
i: [1,1253,0,9,2,0,0,1] d14_gate_1
s: 1253
i: [1,1254,0,9,2,0,0,1] d14_gate_1
s: 1254
i: [1,1256,0,9,2,0,0,1] d17_gate_1
s: 1256
i: [1,1257,0,9,2,0,0,1] d17_gate_1
s: 1257
i: [1,1258,0,9,2,0,0,1] d17_gate_1
s: 1258
i: [1,1259,0,9,2,0,0,1] d17_gate_1
s: 1259
i: [1,1300,0,9,3,0,0,2] fc35_gate_1
s: 1300
i: [1,1301,0,9,3,0,0,2] fc33_gate_1
s: 1301
i: [1,1302,0,9,3,0,0,2] fc34_gate_1
s: 1302
i: [1,1307,0,9,3,0,0,2] fc36_gate_1
s: 1307
i: [1,1308,0,9,3,0,0,2] fc37_gate_1
s: 1308
i: [1,1309,0,9,3,0,0,2] fc38_gate_1
s: 1309
i: [1,1241,0,10,3,0,0,1] d34_gate_1
s: 1241
d: [2,2016,1,7,3,1,44,1656,1241]
d: [2,2024,1,5,2,2,21,2016,1721]
d: [2,2025,1,3,1,3,21,2024,1886]
a: [3,2029,1,3,1,2025]
s: 2029
i: [1,1242,0,10,3,0,0,1] d34_gate_1
s: 1242
i: [1,1263,0,10,2,0,0,1] d18_gate_1
s: 1263
i: [1,1264,0,10,2,0,0,1] d18_gate_1
s: 1264
d: [2,2052,1,4,2,1,44,1392,1264]
d: [2,2053,1,2,1,2,21,2052,1687]
d: [2,2054,1,0,0,3,21,2053,1610]
e: 2054
r: 1392
r: 1397
a: [3,1396,0,2,1,210]
s: 1396
a: [3,1393,0,7,1,210]
s: 1393
i: [1,1265,0,10,2,0,0,1] d18_gate_1
s: 1265
i: [1,1266,0,10,2,0,0,1] d18_gate_1
s: 1266
i: [1,1267,0,10,2,0,0,1] d18_gate_1
s: 1267
i: [1,1269,0,10,2,0,0,1] d21_gate_1
s: 1269
i: [1,1270,0,10,2,0,0,1] d21_gate_1
s: 1270
i: [1,1271,0,10,2,0,0,1] d21_gate_1
s: 1271
i: [1,1272,0,10,2,0,0,1] d21_gate_1
s: 1272
i: [1,1273,0,10,2,0,0,1] d21_gate_1
s: 1273
i: [1,1295,0,11,4,0,0,1] fc17_gate_1
s: 1295
d: [2,2010,1,10,3,0,44,1241,1278]
s: 2010
i: [1,1296,0,11,4,0,0,1] fc32_gate_1
s: 1296
d: [2,2011,1,10,3,0,44,1241,1278]
s: 2011
i: [1,1297,0,11,4,0,0,1] fc31_gate_1
s: 1297
i: [1,1298,0,11,4,0,0,1] fc29_gate_1
s: 1298
d: [2,2014,1,10,3,0,44,1278,1241]
s: 2014
i: [1,1299,0,11,4,0,0,1] fc30_gate_1
s: 1299
d: [2,2015,1,10,3,0,44,1278,1241]
s: 2015
i: [1,1303,0,11,4,0,0,1] fc26_gate_1
s: 1303
i: [1,1304,0,11,4,0,0,1] fc27_gate_1
s: 1304
i: [1,1305,0,11,4,0,0,1] fc28_gate_1
s: 1305
i: [1,1306,0,11,4,0,0,1] fc25_gate_1
s: 1306
i: [1,1246,0,13,4,0,0,1] d13_gate_1
s: 1246
d: [2,2107,1,9,4,1,44,1485,1246]
d: [2,2111,1,6,3,2,21,2107,2029]
d: [2,2112,1,4,2,3,21,2111,1650]
d: [2,2113,1,2,1,4,21,2112,1498]
d: [2,2114,1,0,0,5,21,2113,1508]
e: 2114
r: 1505
r: 1508
r: 1677
r: 1680
r: 1495
r: 1498
a: [3,1499,0,2,1,188]
s: 1499
s: 1508
s: 1680
s: 1505
s: 1677
a: [3,1494,0,5,1,188]
s: 1494
d: [2,2121,1,0,0,2,52,1560,1494,1293]
e: 2121
d: [2,2123,1,2,1,1,40,1494,1293]
d: [2,2126,1,0,0,2,21,2123,1560]
e: 2126
r: 1485
r: 1488
r: 1494
r: 1499
r: 1335
r: 1792
s: 1334
a: [3,1489,0,2,1,190]
s: 1489
s: 1498
s: 1791
a: [3,1484,0,5,1,190]
s: 1484
s: 1495
i: [1,1260,0,14,1,0,0,2] d36_gate_1
s: 1260
i: [1,1261,0,14,1,0,0,1] d35_gate_1
s: 1261
d: [2,2141,1,11,3,0,44,1301,1260]
s: 2141
i: [1,1250,0,16,5,0,0,1] d14_gate_1
s: 1250
i: [1,1255,0,16,5,0,0,1] d17_gate_1
s: 1255
d: [2,2143,1,11,3,0,44,1307,1260]
s: 2143
i: [1,1274,0,18,1,0,0,2] d38_gate_1
s: 1274
d: [2,2178,1,13,3,0,44,1301,1274]
s: 2178
i: [1,1275,0,18,1,0,0,1] d37_gate_1
s: 1275
i: [1,1262,0,19,6,0,0,1] d18_gate_1
s: 1262
d: [2,2180,1,13,3,0,44,1307,1274]
s: 2180
i: [1,1268,0,19,6,0,0,1] d21_gate_1
s: 1268
d: [2,2140,1,14,3,0,44,1300,1260]
s: 2140
i: [1,1276,0,22,1,0,0,1] d40_gate_1
s: 1276
i: [1,1277,0,22,1,0,0,1] d39_gate_1
s: 1277
d: [2,2142,1,14,3,0,44,1302,1260]
s: 2142
d: [2,2144,1,14,3,0,44,1308,1260]
s: 2144
d: [2,2145,1,14,3,0,44,1309,1260]
s: 2145
d: [2,2218,1,15,3,0,44,1301,1276]
s: 2218
d: [2,2220,1,15,3,0,44,1307,1276]
s: 2220
d: [2,2147,1,16,4,0,44,1296,1261]
s: 2147
d: [2,2148,1,16,4,0,44,1297,1261]
s: 2148
d: [2,2149,1,16,4,0,44,1298,1261]
s: 2149
d: [2,2150,1,16,4,0,44,1299,1261]
s: 2150
d: [2,2151,1,16,4,0,44,1303,1261]
s: 2151
d: [2,2152,1,16,4,0,44,1304,1261]
s: 2152
d: [2,2153,1,16,4,0,44,1305,1261]
s: 2153
d: [2,2259,2,9,3,1,40,2153,1323]
d: [2,2265,2,7,2,2,21,2259,1725]
d: [2,2266,2,5,1,3,21,2265,1760]
a: [3,2270,2,5,1,2266]
s: 2270
d: [2,2154,1,16,4,0,44,1306,1261]
s: 2154
d: [2,2177,1,18,3,0,44,1300,1274]
s: 2177
d: [2,2179,1,18,3,0,44,1302,1274]
s: 2179
d: [2,2181,1,18,3,0,44,1308,1274]
s: 2181
d: [2,2182,1,18,3,0,44,1309,1274]
s: 2182
d: [2,2139,1,19,1,0,44,1260,1260]
s: 2139
d: [2,2290,1,18,1,0,22,2139,1274]
s: 2290
d: [2,2146,1,19,1,0,44,1261,1260]
s: 2146
d: [2,2308,1,18,1,0,22,2146,1275]
s: 2308
d: [2,2183,1,20,4,0,44,1296,1275]
s: 2183
r: 1331
a: [3,1330,0,9,1,1236]
s: 1330
d: [2,2184,1,20,4,0,44,1297,1275]
s: 2184
d: [2,2185,1,20,4,0,44,1298,1275]
s: 2185
d: [2,2186,1,20,4,0,44,1299,1275]
s: 2186
d: [2,2328,2,11,3,1,40,2183,1331]
d: [2,2329,2,9,2,2,21,2328,1881]
d: [2,2330,2,7,1,3,21,2329,1877]
a: [3,2333,2,7,1,2330]
s: 2333
d: [2,2187,1,20,4,0,44,1303,1275]
s: 2187
d: [2,2188,1,20,4,0,44,1304,1275]
s: 2188
d: [2,2189,1,20,4,0,44,1305,1275]
s: 2189
d: [2,2190,1,20,4,0,44,1306,1275]
s: 2190
d: [2,2312,2,21,4,0,44,2147,2308]
d: [2,2320,2,20,4,0,22,2312,1260]
s: 2320
d: [2,2217,1,22,3,0,44,1300,1276]
s: 2217
d: [2,2219,1,22,3,0,44,1302,1276]
s: 2219
d: [2,2374,2,11,2,1,40,2217,1347]
d: [2,2376,2,9,1,2,21,2374,1756]
a: [3,2380,2,9,1,2376]
s: 2380
d: [2,2384,3,0,0,3,52,1877,2333,2380,2179]
e: 2384
d: [2,2385,3,0,0,3,52,1881,2333,2380,2177]
e: 2385
d: [2,2386,3,0,0,3,52,1877,1881,2380,2178]
e: 2386
d: [2,2387,3,9,2,1,40,2380,2179]
d: [2,2393,3,2,1,2,21,2387,2333]
d: [2,2394,3,0,0,3,21,2393,1877]
e: 2394
d: [2,2388,3,9,2,1,40,2380,2177]
d: [2,2396,3,2,1,2,21,2388,2333]
d: [2,2397,3,0,0,3,21,2396,1881]
e: 2397
d: [2,2389,3,4,2,1,40,2380,2178]
d: [2,2399,3,2,1,2,21,2389,1881]
d: [2,2400,3,0,0,3,21,2399,1877]
e: 2400
r: 1347
r: 2380
a: [3,1338,0,11,1,1236]
s: 1338
a: [3,1346,0,11,1,1236]
s: 1346
d: [2,2221,1,22,3,0,44,1308,1276]
s: 2221
d: [2,2222,1,22,3,0,44,1309,1276]
s: 2222
a: [3,2379,2,9,1,2376]
s: 2379
d: [2,2175,1,23,1,0,44,1260,1274]
s: 2175
d: [2,2313,2,21,4,0,44,2148,2308]
d: [2,2321,2,20,4,0,22,2313,1260]
s: 2321
d: [2,2416,1,22,1,0,22,2175,1276]
s: 2416
d: [2,2176,1,23,1,0,44,1261,1274]
s: 2176
d: [2,2314,2,21,4,0,44,2149,2308]
d: [2,2322,2,20,4,0,22,2314,1260]
s: 2322
d: [2,2431,1,22,1,0,22,2176,1277]
s: 2431
d: [2,2315,2,21,4,0,44,2150,2308]
d: [2,2323,2,20,4,0,22,2315,1260]
s: 2323
d: [2,2223,1,24,4,0,44,1296,1277]
s: 2223
d: [2,2224,1,24,4,0,44,1297,1277]
s: 2224
d: [2,2450,3,0,0,4,52,1756,1731,2379,1338,2224]
e: 2450
d: [2,2451,2,13,3,1,40,2224,1338]
d: [2,2455,2,11,2,2,21,2451,1731]
d: [2,2456,2,2,1,3,21,2455,2379]
d: [2,2457,2,0,0,4,21,2456,1756]
e: 2457
r: 1330
r: 1338
r: 2333
a: [3,2334,2,7,1,2330]
s: 2334
d: [2,2459,3,0,0,3,52,1725,2270,2334,2142]
e: 2459
d: [2,2460,3,7,2,1,40,2334,2142]
d: [2,2464,3,2,1,2,21,2460,2270]
d: [2,2465,3,0,0,3,21,2464,1725]
e: 2465
r: 1607
r: 1610
r: 1881
r: 2334
r: 1566
r: 1571
r: 1656
r: 1661
r: 2029
r: 1315
s: 1571
s: 1610
s: 1661
s: 1566
s: 1881
s: 1607
s: 2029
s: 1656
a: [3,1314,0,5,1,1236]
s: 1314
d: [2,2509,2,0,0,4,52,1886,1702,1721,1314,1297]
e: 2509
d: [2,2513,1,6,3,1,40,1314,1297]
d: [2,2515,1,4,2,2,21,2513,1721]
d: [2,2516,1,2,1,3,21,2515,1702]
d: [2,2517,1,0,0,4,21,2516,1886]
e: 2517
r: 1666
r: 1671
r: 1721
r: 1314
r: 1597
r: 1600
r: 1607
r: 1610
r: 1881
r: 1350
r: 1386
a: [3,1351,0,2,1,1236]
s: 1351
s: 1600
s: 1610
a: [3,1670,0,2,1,152]
s: 1670
s: 1597
s: 1721
s: 1607
s: 1881
a: [3,1667,0,4,1,152]
s: 1667
d: [2,2537,2,0,0,2,52,1721,1667,1284]
e: 2537
d: [2,2539,1,2,1,1,40,1667,1284]
d: [2,2541,1,0,0,2,21,2539,1721]
e: 2541
r: 1667
r: 1670
s: 1671
s: 1666
s: 1315
a: [3,1387,0,7,1,212]
s: 1387
d: [2,2548,1,13,6,1,44,1387,1268]
d: [2,2554,1,10,5,2,21,2548,2029]
d: [2,2555,1,8,4,3,21,2554,1396]
d: [2,2556,1,6,3,4,21,2555,1406]
d: [2,2557,1,4,2,5,21,2556,1416]
d: [2,2558,1,2,1,6,21,2557,1426]
d: [2,2559,1,0,0,7,21,2558,1590]
e: 2559
r: 1587
r: 1590
s: 1591
s: 1586
d: [2,2562,3,0,0,2,52,1756,1586,1806]
e: 2562
d: [2,2563,2,2,1,1,40,1586,1806]
d: [2,2566,2,0,0,2,21,2563,1756]
e: 2566
r: 1403
r: 1406
r: 1586
r: 1591
a: [3,1407,0,2,1,208]
s: 1407
s: 1590
s: 1587
a: [3,1402,0,6,1,208]
s: 1402
d: [2,2574,1,4,2,1,44,1402,1253]
d: [2,2577,1,2,1,2,21,2574,1687]
d: [2,2578,1,0,0,3,21,2577,1610]
e: 2578
d: [2,2576,1,4,2,1,44,1402,1251]
d: [2,2580,1,2,1,2,21,2576,1687]
d: [2,2581,1,0,0,3,21,2580,1680]
e: 2581
r: 1402
r: 1407
r: 1423
r: 1426
s: 1406
a: [3,1427,0,2,1,204]
s: 1427
a: [3,1422,0,4,1,204]
s: 1422
d: [2,2585,1,0,0,2,52,1620,1422,1285]
e: 2585
d: [2,2587,1,2,1,1,40,1422,1285]
d: [2,2589,1,0,0,2,21,2587,1620]
e: 2589
r: 1617
r: 1620
r: 1438
r: 1443
r: 1323
r: 1351
r: 1387
r: 1422
r: 1427
s: 1350
s: 1426
s: 1443
s: 1620
s: 1423
s: 1617
s: 1403
s: 1438
a: [3,1322,0,7,1,1236]
s: 1322
d: [2,2608,3,0,0,4,52,1725,1760,2270,1322,2150]
e: 2608
d: [2,2610,2,9,3,1,40,1322,2150]
d: [2,2614,2,7,2,2,21,2610,1760]
d: [2,2615,2,2,1,3,21,2614,2270]
d: [2,2616,2,0,0,4,21,2615,1725]
e: 2616
r: 1322
r: 1731
r: 1756
s: 1386
s: 1730
s: 1330
s: 1338
a: [3,1755,2,2,1,1749]
s: 1755
d: [2,2225,1,24,4,0,44,1298,1277]
s: 2225
s: 2333
d: [2,2226,1,24,4,0,44,1299,1277]
s: 2226
d: [2,2645,3,0,0,4,52,1730,1755,2379,1338,2226]
e: 2645
d: [2,2647,2,13,3,1,40,2226,1338]
d: [2,2650,2,11,2,2,21,2647,1755]
d: [2,2651,2,2,1,3,21,2650,2379]
d: [2,2652,2,0,0,4,21,2651,1730]
e: 2652
r: 1438
r: 1443
r: 1566
r: 1571
r: 1576
r: 1581
r: 1755
a: [3,1442,0,2,1,200]
s: 1442
a: [3,1570,0,2,1,172]
s: 1570
s: 1580
s: 1756
a: [3,1567,0,4,1,172]
s: 1567
s: 1577
a: [3,1439,0,6,1,200]
s: 1439
d: [2,2665,1,11,5,1,44,1439,1255]
d: [2,2670,1,8,4,2,21,2665,2029]
d: [2,2671,1,6,3,3,21,2670,1620]
d: [2,2672,1,4,2,4,21,2671,1452]
d: [2,2673,1,2,1,5,21,2672,1462]
d: [2,2674,1,0,0,6,21,2673,1472]
e: 2674
r: 1449
r: 1452
a: [3,1453,0,2,1,198]
s: 1453
a: [3,1448,0,6,1,198]
s: 1448
d: [2,2680,1,4,2,1,44,1448,1254]
d: [2,2684,1,2,1,2,21,2680,1687]
d: [2,2685,1,0,0,3,21,2684,1560]
e: 2685
d: [2,2683,1,4,2,1,44,1448,1251]
d: [2,2687,1,2,1,2,21,2683,1687]
d: [2,2688,1,0,0,3,21,2687,1610]
e: 2688
r: 1448
r: 1453
r: 1459
r: 1462
s: 1452
a: [3,1463,0,2,1,196]
s: 1463
a: [3,1458,0,5,1,196]
s: 1458
d: [2,2692,1,0,0,2,52,1680,1458,1294]
e: 2692
d: [2,2693,1,0,0,2,52,1610,1458,1292]
e: 2693
d: [2,2694,1,2,1,1,40,1458,1294]
d: [2,2699,1,0,0,2,21,2694,1680]
e: 2699
d: [2,2696,1,2,1,1,40,1458,1292]
d: [2,2701,1,0,0,2,21,2696,1610]
e: 2701
r: 1525
r: 1607
r: 1610
r: 1677
r: 1680
r: 1810
r: 1877
r: 1327
r: 1330
r: 1484
r: 1489
r: 1505
r: 1508
r: 1597
r: 1600
r: 1787
s: 1326
s: 1488
s: 1509
s: 1524
s: 1601
s: 1611
s: 1681
s: 1504
s: 1596
a: [3,1878,1,2,1,1866]
s: 1878
s: 1606
s: 1676
d: [2,2731,2,0,0,2,52,1886,1676,1806]
e: 2731
d: [2,2732,2,2,1,1,40,1676,1806]
d: [2,2735,2,0,0,2,21,2732,1886]
e: 2735
r: 1886
r: 1458
r: 1463
r: 1495
r: 1498
r: 1504
r: 1509
r: 1557
r: 1560
r: 1617
r: 1620
r: 1676
r: 1681
r: 1702
r: 1721
r: 1315
r: 1534
r: 1546
r: 1551
r: 1769
r: 2270
r: 1319
r: 1781
r: 2333
r: 1334
r: 1488
r: 1791
s: 1319
s: 1335
s: 1462
s: 1489
s: 1498
s: 1508
s: 1535
s: 1551
s: 1560
a: [3,1621,0,2,1,162]
s: 1621
s: 1680
s: 1546
s: 1702
s: 1557
a: [3,1720,1,2,1,1714]
s: 1720
s: 1505
a: [3,1616,0,4,1,162]
s: 1616
d: [2,2779,2,0,0,2,52,1881,1616,1806]
e: 2779
d: [2,2780,2,2,1,1,40,1616,1806]
d: [2,2783,2,0,0,2,21,2780,1881]
e: 2783
r: 1343
r: 1439
r: 1442
r: 1616
r: 1621
r: 1797
s: 1342
s: 1443
s: 1620
s: 1617
s: 1677
a: [3,1887,1,2,1,1864]
s: 1887
s: 1768
d: [2,2789,2,0,0,2,52,1560,1768,1764]
e: 2789
d: [2,2790,2,2,1,1,40,1768,1764]
d: [2,2793,2,0,0,2,21,2790,1560]
e: 2793
r: 1535
r: 1768
s: 1534
s: 1769
s: 1781
s: 1786
s: 1792
s: 1796
a: [3,1811,0,4,1,1805]
s: 1811
d: [2,2806,1,0,0,3,52,1680,1534,1811,1290]
e: 2806
d: [2,2807,1,4,2,1,40,1811,1290]
d: [2,2809,1,2,1,2,21,2807,1534]
d: [2,2810,1,0,0,3,21,2809,1680]
e: 2810
r: 1326
r: 1524
r: 1786
r: 1811
s: 1327
s: 1525
s: 1787
s: 1810
s: 1314
d: [2,2820,2,0,0,4,52,1720,1702,1887,1314,1299]
e: 2820
d: [2,2822,1,6,3,1,40,1314,1299]
d: [2,2826,1,4,2,2,21,2822,1887]
d: [2,2827,1,2,1,3,21,2826,1702]
d: [2,2828,1,0,0,4,21,2827,1720]
e: 2828
r: 1314
r: 1350
r: 1386
r: 1393
r: 1396
s: 1351
s: 1397
s: 1459
s: 1484
s: 1495
s: 1438
s: 2270
s: 1449
s: 1387
s: 1392
d: [2,2856,1,4,2,1,44,1392,1267]
d: [2,2861,1,2,1,2,21,2856,1687]
d: [2,2862,1,0,0,3,21,2861,1560]
e: 2862
d: [2,2860,1,4,2,1,44,1392,1263]
d: [2,2864,1,2,1,2,21,2860,1687]
d: [2,2865,1,0,0,3,21,2864,1580]
e: 2865
r: 1327
r: 1351
r: 1387
r: 1392
r: 1397
r: 1626
r: 1631
r: 1787
s: 1326
s: 1350
s: 1396
a: [3,1630,0,2,1,160]
s: 1630
a: [3,1627,0,4,1,160]
s: 1627
d: [2,2869,1,7,3,1,44,1627,1242]
d: [2,2873,1,4,2,2,21,2869,2029]
d: [2,2874,1,2,1,3,21,2873,1760]
d: [2,2875,1,0,0,4,21,2874,1725]
e: 2875
r: 1760
r: 1636
r: 1641
s: 1640
s: 1761
s: 1637
s: 1786
s: 1323
d: [2,2885,3,0,0,4,52,1725,1761,2270,1323,2147]
e: 2885
d: [2,2889,2,9,3,1,40,1323,2147]
d: [2,2891,2,4,2,2,21,2889,2270]
d: [2,2892,2,2,1,3,21,2891,1761]
d: [2,2893,2,0,0,4,21,2892,1725]
e: 2893
r: 2270
s: 1386
a: [3,2269,2,5,1,2266]
s: 2269
d: [2,2904,3,0,0,3,52,1720,1702,2269,1302]
e: 2904
d: [2,2905,3,4,2,1,40,2269,1302]
d: [2,2909,3,2,1,2,21,2905,1702]
d: [2,2910,3,0,0,3,21,2909,1720]
e: 2910
r: 1323
r: 1326
r: 1786
r: 2269
s: 1327
s: 1787
d: [2,2914,1,0,0,3,52,1630,1525,1787,1289]
e: 2914
d: [2,2916,1,4,2,1,40,1787,1289]
d: [2,2918,1,2,1,2,21,2916,1525]
d: [2,2919,1,0,0,3,21,2918,1630]
e: 2919
r: 1319
r: 1327
r: 1342
r: 1438
r: 1443
r: 1449
r: 1452
r: 1546
r: 1551
r: 1557
r: 1560
r: 1666
r: 1671
r: 1702
r: 1781
r: 1787
r: 1796
s: 1318
s: 1326
s: 1343
s: 1442
s: 1453
s: 1550
s: 1561
s: 1670
s: 1547
a: [3,1703,1,2,1,1699]
s: 1703
s: 1556
s: 1667
d: [2,2930,2,0,0,2,52,1887,1667,1747]
e: 2930
d: [2,2931,2,2,1,1,40,1667,1747]
d: [2,2934,2,0,0,2,21,2931,1887]
e: 2934
r: 1525
r: 1534
r: 1667
r: 1670
r: 1769
r: 1810
s: 1524
s: 1535
s: 1671
s: 1666
s: 1768
s: 1782
s: 1786
d: [2,2945,1,0,0,3,52,1630,1524,1786,1291]
e: 2945
d: [2,2946,1,4,2,1,40,1786,1291]
d: [2,2949,1,2,1,2,21,2946,1524]
d: [2,2950,1,0,0,3,21,2949,1630]
e: 2950
r: 1326
r: 1786
s: 1327
s: 1787
s: 1797
s: 1811
s: 1314
s: 2269
s: 1439
s: 1448
d: [2,2974,1,4,2,1,44,1448,1252]
d: [2,2976,1,2,1,2,21,2974,1687]
d: [2,2977,1,0,0,3,21,2976,1640]
e: 2977
r: 1439
r: 1442
r: 1448
r: 1453
r: 1567
r: 1570
r: 1577
r: 1580
r: 1756
r: 2379
r: 1346
s: 1443
s: 1452
s: 1571
s: 1581
s: 1566
s: 1755
s: 1576
s: 1438
s: 1449
s: 1322
d: [2,3004,3,0,0,4,52,1725,1761,2269,1322,2149]
e: 3004
d: [2,3007,2,9,3,1,40,1322,2149]
d: [2,3010,2,4,2,2,21,3007,2269]
d: [2,3011,2,2,1,3,21,3010,1761]
d: [2,3012,2,0,0,4,21,3011,1725]
e: 3012
r: 1322
s: 1393
s: 1331
r: 1881
r: 1335
r: 1343
r: 1438
r: 1443
r: 1596
r: 1601
r: 1617
r: 1620
r: 1792
r: 1797
r: 1423
r: 1426
r: 1350
r: 1386
s: 1334
s: 1342
s: 1351
s: 1427
s: 1442
s: 1600
s: 1621
s: 1422
a: [3,1882,1,2,1,1866]
s: 1882
s: 1597
s: 1616
s: 1791
s: 1796
s: 1439
s: 1387
s: 2334
d: [2,3067,3,0,0,3,52,1725,1761,2334,2141]
e: 3067
d: [2,3070,3,4,2,1,40,2334,2141]
d: [2,3072,3,2,1,2,21,3070,1761]
d: [2,3073,3,0,0,3,21,3072,1725]
e: 3073
r: 1327
r: 1627
r: 1630
r: 1637
r: 1640
r: 1761
r: 1787
r: 1393
r: 1396
r: 1449
r: 1452
r: 1495
r: 1498
r: 1318
r: 1334
r: 1484
r: 1489
r: 1656
r: 1661
r: 1782
r: 1791
s: 1319
s: 1326
s: 1335
s: 1397
s: 1453
s: 1488
s: 1499
s: 1631
s: 1641
a: [3,1660,0,2,1,154]
s: 1660
s: 1626
s: 1636
s: 1760
a: [3,1657,0,4,1,154]
s: 1657
d: [2,3109,1,7,3,1,44,1657,1242]
d: [2,3113,1,4,2,2,21,3109,2029]
d: [2,3114,1,2,1,3,21,3113,1887]
d: [2,3115,1,0,0,4,21,3114,1720]
e: 3115
r: 1319
r: 1326
r: 1626
r: 1631
r: 1657
r: 1660
r: 1725
r: 1469
r: 1472
r: 1647
r: 1650
r: 1413
r: 1416
s: 1318
s: 1327
s: 1417
s: 1473
s: 1630
s: 1651
s: 1661
s: 1468
s: 1726
s: 1627
s: 1646
s: 1656
s: 1782
s: 1787
s: 1792
s: 1412
s: 1485
s: 1494
s: 1448
s: 1322
d: [2,3159,3,0,0,4,52,1726,1760,2269,1322,2154]
e: 3159
d: [2,3160,2,9,3,1,40,1322,2154]
d: [2,3165,2,4,2,2,21,3160,2269]
d: [2,3166,2,2,1,3,21,3165,1760]
d: [2,3167,2,0,0,4,21,3166,1726]
e: 3167
r: 1322
r: 1342
r: 1566
r: 1571
r: 1796
s: 1343
s: 1570
s: 1567
d: [2,3171,1,7,3,1,44,1567,1242]
d: [2,3175,1,4,2,2,21,3171,2029]
d: [2,3176,1,2,1,3,21,3175,1755]
d: [2,3177,1,0,0,4,21,3176,1730]
e: 3177
r: 1314
r: 1318
r: 1327
r: 1343
r: 1397
r: 1448
r: 1453
r: 1494
r: 1499
r: 1524
r: 1535
r: 1547
r: 1550
r: 1556
r: 1561
r: 1567
r: 1570
r: 1703
r: 1768
r: 1782
r: 1787
r: 1811
r: 2269
s: 1319
s: 1326
s: 1342
s: 1396
s: 1452
s: 1498
s: 1525
s: 1534
s: 1551
s: 1560
s: 1571
s: 1546
s: 1702
s: 1557
s: 1566
s: 1769
s: 1781
s: 1786
s: 1796
s: 1810
s: 1495
s: 1449
s: 1323
r: 1323
r: 1335
r: 1485
r: 1488
r: 1566
r: 1571
r: 1730
r: 1755
r: 1792
s: 1334
s: 1489
s: 1571
s: 1566
s: 1731
s: 1791
s: 1484
d: [2,3236,1,4,2,1,44,1484,1249]
d: [2,3239,1,2,1,2,21,3236,1687]
d: [2,3240,1,0,0,3,21,3239,1651]
e: 3240
r: 1334
r: 1342
r: 1484
r: 1489
r: 1566
r: 1571
r: 1791
r: 1796
s: 1335
s: 1343
s: 1488
s: 1570
s: 1567
r: 1319
r: 1326
r: 1343
r: 1396
r: 1412
r: 1417
r: 1449
r: 1452
r: 1468
r: 1473
r: 1495
r: 1498
r: 1525
r: 1534
r: 1546
r: 1551
r: 1557
r: 1560
r: 1567
r: 1570
r: 1646
r: 1651
r: 1702
r: 1760
r: 1769
r: 1781
r: 1786
r: 1810
r: 1587
r: 1590
r: 1627
r: 1630
s: 1318
s: 1326
s: 1343
s: 1397
s: 1416
s: 1453
s: 1472
s: 1499
s: 1524
s: 1535
s: 1550
s: 1561
s: 1570
s: 1591
s: 1631
s: 1650
s: 1547
s: 1703
s: 1556
s: 1469
s: 1755
s: 1567
s: 1761
s: 1586
s: 1626
s: 1647
s: 1768
s: 1782
s: 1786
s: 1792
s: 1797
s: 1811
s: 1314
s: 1413
s: 1485
s: 1494
s: 1448
s: 2269
s: 1322
s: 1392
s: 2380
s: 1347
d: [2,2316,2,21,4,0,44,2151,2308]
d: [2,2324,2,20,4,0,22,2316,1260]
s: 2324
d: [2,2227,1,24,4,0,44,1303,1277]
s: 2227
d: [2,2228,1,24,4,0,44,1304,1277]
s: 2228
d: [2,2317,2,21,4,0,44,2152,2308]
d: [2,2325,2,20,4,0,22,2317,1260]
s: 2325
d: [2,2229,1,24,4,0,44,1305,1277]
s: 2229
d: [2,2318,2,21,4,0,44,2153,2308]
d: [2,2326,2,20,4,0,22,2318,1260]
s: 2326
d: [2,2230,1,24,4,0,44,1306,1277]
s: 2230
d: [2,3347,3,0,0,4,52,1731,1755,2380,1338,2230]
e: 3347
d: [2,3352,2,13,3,1,40,2230,1338]
d: [2,3357,2,4,2,2,21,3352,2380]
d: [2,3358,2,2,1,3,21,3357,1755]
d: [2,3359,2,0,0,4,21,3358,1731]
e: 3359
r: 1422
r: 1427
r: 1597
r: 1600
r: 1616
r: 1621
r: 1882
r: 2334
r: 2380
r: 1335
r: 1792
s: 1334
s: 1426
s: 1601
s: 1620
s: 1423
s: 1596
s: 1881
s: 1617
s: 1791
d: [2,2215,1,27,1,0,44,1260,1276]
s: 2215
d: [2,2216,1,27,1,0,44,1261,1276]
s: 2216
s: 2333
d: [2,3413,3,0,0,3,52,1726,2269,2333,2144]
e: 3413
d: [2,3415,3,7,2,1,40,2333,2144]
d: [2,3418,3,2,1,2,21,3415,2269]
d: [2,3419,3,0,0,3,21,3418,1726]
e: 3419
r: 1314
r: 1318
r: 1326
r: 1334
r: 1343
r: 1392
r: 1397
r: 1439
r: 1442
r: 1448
r: 1453
r: 1485
r: 1488
r: 1494
r: 1499
r: 1524
r: 1535
r: 1547
r: 1550
r: 1556
r: 1561
r: 1703
r: 1768
r: 1782
r: 1786
r: 1791
r: 1797
r: 1811
r: 2269
r: 1322
r: 1351
r: 1387
s: 1319
s: 1327
s: 1335
s: 1342
s: 1350
s: 1396
s: 1443
s: 1452
s: 1489
s: 1498
s: 1525
s: 1534
s: 1551
s: 1560
s: 1702
s: 1546
s: 1557
s: 1769
s: 1781
s: 1787
s: 1792
s: 1796
s: 1810
s: 1484
s: 1495
s: 1438
s: 2270
s: 1449
s: 1323
s: 1386
d: [2,3478,1,4,2,1,44,1386,1269]
d: [2,3479,1,2,1,2,21,3478,1687]
d: [2,3480,1,0,0,3,21,3479,1591]
e: 3480
r: 1323
r: 1327
r: 1335
r: 1342
r: 1350
r: 1386
r: 1413
r: 1416
r: 1438
r: 1443
r: 1469
r: 1472
r: 1484
r: 1489
r: 1626
r: 1631
r: 1647
r: 1650
r: 1761
r: 1787
r: 1792
r: 1796
r: 1449
r: 1452
r: 1495
r: 1498
r: 1525
r: 1534
r: 1546
r: 1551
r: 1557
r: 1560
r: 1656
r: 1661
r: 1666
r: 1671
r: 1702
r: 1726
r: 1769
r: 1810
r: 1887
s: 1326
s: 1334
s: 1343
s: 1351
s: 1416
s: 1442
s: 1453
s: 1472
s: 1488
s: 1499
s: 1524
s: 1535
s: 1550
s: 1561
s: 1631
s: 1650
s: 1660
s: 1670
s: 1547
s: 1703
s: 1556
s: 1725
s: 1469
s: 1626
s: 1886
s: 1647
s: 1760
s: 1657
s: 1667
s: 1768
d: [2,3516,1,0,0,2,52,1670,1768,1285]
e: 3516
d: [2,3518,1,2,1,1,40,1768,1285]
d: [2,3520,1,0,0,2,21,3518,1670]
e: 3520
r: 1319
r: 1326
r: 1334
r: 1488
r: 1499
r: 1524
r: 1535
r: 1768
r: 1781
s: 1318
s: 1327
s: 1335
s: 1489
s: 1498
s: 1525
s: 1534
s: 1769
s: 1782
d: [2,3533,1,0,0,3,52,1660,1550,1782,1291]
e: 3533
d: [2,3534,1,4,2,1,40,1782,1291]
d: [2,3537,1,2,1,2,21,3534,1550]
d: [2,3538,1,0,0,3,21,3537,1660]
e: 3538
r: 1318
r: 1782
s: 1319
s: 1781
s: 1787
s: 1792
s: 1797
s: 1810
s: 1413
s: 1484
s: 1495
s: 1439
s: 1448
d: [2,3565,1,4,2,1,44,1448,1253]
d: [2,3568,1,2,1,2,21,3565,1687]
d: [2,3569,1,0,0,3,21,3568,1670]
e: 3569
r: 1423
r: 1426
r: 1448
r: 1453
r: 1617
r: 1620
r: 1878
r: 1881
s: 1427
s: 1452
s: 1621
s: 1422
s: 1877
s: 1616
d: [2,3585,2,0,0,2,52,1877,1616,1283]
e: 3585
d: [2,3587,1,2,1,1,40,1616,1283]
d: [2,3589,1,0,0,2,21,3587,1877]
e: 3589
r: 1343
r: 1422
r: 1427
r: 1439
r: 1442
r: 1616
r: 1621
r: 1725
r: 1760
r: 1797
r: 1877
s: 1342
s: 1426
s: 1443
s: 1620
s: 1423
s: 1726
s: 1617
s: 1796
s: 1878
s: 1314
d: [2,3602,2,0,0,4,52,1720,1886,1703,1314,1298]
e: 3602
d: [2,3605,1,6,3,1,40,1314,1298]
d: [2,3608,1,4,2,2,21,3605,1703]
d: [2,3609,1,2,1,3,21,3608,1886]
d: [2,3610,1,0,0,4,21,3609,1720]
e: 3610
r: 1314
r: 1335
r: 1484
r: 1489
r: 1495
r: 1498
r: 1792
s: 1334
s: 1488
s: 1499
s: 1881
s: 1791
s: 1761
s: 1485
s: 1494
d: [2,3620,1,0,0,2,52,1670,1494,1292]
e: 3620
d: [2,3623,1,2,1,1,40,1494,1292]
d: [2,3625,1,0,0,2,21,3623,1670]
e: 3625
r: 1319
r: 1327
r: 1334
r: 1494
r: 1499
r: 1547
r: 1550
r: 1556
r: 1561
r: 1596
r: 1601
r: 1606
r: 1611
r: 1626
r: 1631
r: 1647
r: 1650
r: 1667
r: 1670
r: 1703
r: 1720
r: 1761
r: 1781
r: 1787
r: 1791
r: 1878
r: 2333
r: 1342
r: 1413
r: 1416
r: 1469
r: 1472
r: 1567
r: 1570
r: 1586
r: 1591
r: 1731
r: 1796
r: 1881
r: 1396
r: 1423
r: 1426
r: 1443
r: 1452
r: 1525
r: 1534
r: 1617
r: 1620
r: 1657
r: 1660
r: 1769
r: 1810
r: 1886
r: 2270
s: 1318
s: 1326
s: 1335
s: 1342
s: 1397
s: 1416
s: 1427
s: 1442
s: 1453
s: 1472
s: 1499
s: 1524
s: 1535
s: 1550
s: 1561
s: 1571
s: 1590
s: 1600
s: 1611
s: 1621
s: 1631
s: 1650
s: 1661
s: 1671
s: 1547
s: 1703
s: 1556
s: 1720
s: 1422
s: 1469
s: 1730
s: 1566
s: 1878
s: 1587
s: 1597
s: 1882
s: 1606
s: 1887
s: 1616
s: 1626
s: 1761
s: 1647
s: 1656
s: 1666
s: 1768
s: 1782
s: 1786
s: 1792
s: 1796
s: 1811
s: 1314
s: 1413
s: 1494
s: 1439
s: 1448
s: 2269
s: 1323
d: [2,3732,3,0,0,4,52,1726,1761,2269,1323,2152]
e: 3732
d: [2,3734,2,9,3,1,40,1323,2152]
d: [2,3738,2,7,2,2,21,3734,1761]
d: [2,3739,2,2,1,3,21,3738,2269]
d: [2,3740,2,0,0,4,21,3739,1726]
e: 3740
r: 1326
r: 1335
r: 1347
r: 1351
r: 1397
r: 1422
r: 1427
r: 1439
r: 1442
r: 1448
r: 1453
r: 1566
r: 1571
r: 1576
r: 1581
r: 1597
r: 1600
r: 1616
r: 1621
r: 1626
r: 1631
r: 1636
r: 1641
r: 1726
r: 1755
r: 1786
r: 1792
r: 1878
r: 1494
r: 1499
r: 1505
r: 1508
r: 1535
r: 1547
r: 1550
r: 1556
r: 1561
r: 1656
r: 1661
r: 1677
r: 1680
r: 1703
r: 1720
r: 1768
s: 1327
s: 1334
s: 1350
s: 1396
s: 1426
s: 1443
s: 1452
s: 1498
s: 1509
s: 1534
s: 1551
s: 1560
s: 1570
s: 1580
s: 1601
s: 1620
s: 1630
s: 1640
s: 1660
s: 1681
s: 1702
s: 1546
s: 1557
s: 1721
s: 1423
s: 1725
s: 1504
d: [2,3770,2,0,0,2,52,1640,1504,1764]
e: 3770
d: [2,3771,2,2,1,1,40,1504,1764]
d: [2,3774,2,0,0,2,21,3771,1640]
e: 3774
r: 1334
r: 1342
r: 1443
r: 1459
r: 1462
r: 1485
r: 1488
r: 1504
r: 1509
r: 1796
s: 1335
s: 1343
s: 1442
s: 1463
s: 1489
s: 1508
s: 1877
s: 1505
s: 1756
s: 1567
s: 1577
s: 1596
s: 1617
s: 1627
s: 1637
s: 1657
s: 1676
s: 1769
s: 1787
s: 1792
s: 1797
s: 1458
d: [2,3811,1,0,0,2,52,1640,1458,1293]
e: 3811
d: [2,3813,1,2,1,1,40,1458,1293]
d: [2,3816,1,0,0,2,21,3813,1640]
e: 3816
r: 1343
r: 1350
r: 1403
r: 1406
r: 1442
r: 1458
r: 1463
r: 1797
s: 1342
s: 1351
s: 1407
s: 1443
s: 1462
s: 1796
s: 1459
s: 1484
s: 1495
s: 1402
d: [2,3829,1,4,2,1,44,1402,1254]
d: [2,3833,1,2,1,2,21,3829,1687]
d: [2,3834,1,0,0,3,21,3833,1640]
e: 3834
d: [2,3831,1,4,2,1,44,1402,1252]
d: [2,3836,1,2,1,2,21,3831,1687]
d: [2,3837,1,0,0,3,21,3836,1580]
e: 3837
r: 1323
r: 1327
r: 1335
r: 1351
r: 1402
r: 1407
r: 1505
r: 1508
r: 1596
r: 1601
r: 1606
r: 1611
r: 1627
r: 1630
r: 1637
r: 1640
r: 1761
r: 1787
r: 1792
r: 1882
r: 1338
r: 1342
r: 1423
r: 1426
r: 1443
r: 1469
r: 1472
r: 1484
r: 1489
r: 1524
r: 1617
r: 1620
r: 1647
r: 1650
r: 1666
r: 1671
r: 1676
r: 1681
r: 1721
r: 1725
r: 1796
r: 1811
r: 1877
r: 1887
r: 2269
s: 1326
s: 1335
s: 1343
s: 1351
s: 1406
s: 1427
s: 1442
s: 1473
s: 1488
s: 1508
s: 1525
s: 1600
s: 1611
s: 1621
s: 1630
s: 1641
s: 1651
s: 1670
s: 1680
s: 1720
s: 1422
d: [2,3862,2,0,0,2,52,1580,1422,1764]
e: 3862
d: [2,3863,2,2,1,1,40,1422,1764]
d: [2,3866,2,0,0,2,21,3863,1580]
e: 3866
r: 1343
r: 1351
r: 1422
r: 1427
r: 1442
r: 1473
r: 1495
r: 1498
r: 1525
r: 1534
r: 1546
r: 1551
r: 1557
r: 1560
r: 1611
r: 1621
r: 1630
r: 1651
r: 1657
r: 1660
r: 1670
r: 1702
r: 1720
r: 1769
s: 1342
s: 1350
s: 1426
s: 1443
s: 1472
s: 1499
s: 1524
s: 1535
s: 1550
s: 1561
s: 1610
s: 1620
s: 1631
s: 1650
s: 1661
s: 1671
s: 1547
s: 1556
s: 1703
s: 1423
s: 1721
s: 1469
s: 1505
s: 1725
s: 1597
s: 1877
s: 1607
s: 1617
s: 1881
s: 1626
s: 1886
s: 1636
s: 1647
s: 1760
s: 1656
s: 1666
s: 1677
s: 1768
s: 1786
s: 1792
s: 1796
s: 1811
s: 1485
s: 1494
s: 1403
s: 2269
s: 1438
s: 1449
s: 1322
s: 1386
s: 1393
a: [3,1339,0,11,1,1236]
s: 1339
r: 1322
r: 1326
r: 1342
r: 1350
r: 1386
r: 1393
r: 1396
r: 1485
r: 1488
r: 1494
r: 1499
r: 1567
r: 1570
r: 1577
r: 1580
r: 1597
r: 1600
r: 1607
r: 1610
r: 1626
r: 1631
r: 1636
r: 1641
r: 1756
r: 1760
r: 1786
r: 1796
r: 1877
s: 1327
s: 1343
s: 1351
s: 1397
s: 1489
s: 1498
s: 1571
s: 1581
s: 1601
s: 1611
s: 1630
s: 1640
s: 1566
s: 1878
s: 1576
s: 1596
s: 1755
s: 1606
s: 1761
s: 1627
s: 1637
s: 1787
s: 1797
s: 1484
s: 1495
s: 1387
s: 1392
d: [2,4023,1,4,2,1,44,1392,1265]
d: [2,4026,1,2,1,2,21,4023,1687]
d: [2,4027,1,0,0,3,21,4026,1640]
e: 4027
r: 1331
r: 1351
r: 1387
r: 1392
r: 1397
r: 1524
r: 1535
r: 1547
r: 1550
r: 1556
r: 1561
r: 1627
r: 1630
r: 1637
r: 1640
r: 1656
r: 1661
r: 1666
r: 1671
r: 1703
r: 1721
r: 1725
r: 1768
r: 1811
r: 1878
r: 1881
r: 2269
s: 1350
s: 1396
s: 1525
s: 1534
s: 1551
s: 1560
s: 1631
s: 1641
s: 1660
s: 1670
s: 1546
s: 1702
s: 1557
s: 1626
s: 1720
s: 1636
s: 1726
s: 1657
s: 1667
s: 1877
s: 1769
s: 1882
s: 1810
s: 1323
s: 2270
s: 1386
s: 1393
s: 1330
r: 1335
r: 1343
r: 1350
r: 1386
r: 1423
r: 1426
r: 1438
r: 1443
r: 1596
r: 1601
r: 1617
r: 1620
r: 1792
r: 1797
r: 1877
s: 1334
s: 1342
s: 1351
s: 1427
s: 1442
s: 1600
s: 1621
s: 1422
s: 1878
s: 1597
s: 1616
s: 1791
s: 1796
s: 1439
s: 1387
s: 1346
s: 2333
s: 2379
d: [2,4132,3,0,0,3,52,1878,1882,2379,2180]
e: 4132
d: [2,4135,3,4,2,1,40,2379,2180]
d: [2,4137,3,2,1,2,21,4135,1882]
d: [2,4138,3,0,0,3,21,4137,1878]
e: 4138
r: 1323
r: 1327
r: 1351
r: 1387
r: 1422
r: 1427
r: 1439
r: 1442
r: 1469
r: 1472
r: 1484
r: 1489
r: 1566
r: 1571
r: 1576
r: 1581
r: 1597
r: 1600
r: 1616
r: 1621
r: 1626
r: 1631
r: 1647
r: 1650
r: 1755
r: 1761
r: 1787
r: 1878
r: 2333
r: 2379
s: 1326
s: 1350
s: 1426
s: 1443
s: 1473
s: 1488
s: 1570
s: 1580
s: 1601
s: 1620
s: 1630
s: 1651
s: 1423
s: 1877
s: 1468
s: 1756
s: 1567
s: 1577
s: 1760
s: 1596
s: 1617
s: 1627
s: 1646
s: 1786
s: 1485
s: 1438
d: [2,4179,1,4,2,1,44,1438,1256]
d: [2,4180,1,2,1,2,21,4179,1687]
d: [2,4181,1,0,0,3,21,4180,1473]
e: 4181
r: 1449
r: 1452
r: 1468
r: 1473
r: 1495
r: 1498
r: 1525
r: 1534
r: 1546
r: 1551
r: 1557
r: 1560
r: 1627
r: 1630
r: 1646
r: 1651
r: 1657
r: 1660
r: 1667
r: 1670
r: 1702
r: 1720
r: 1726
r: 1769
r: 1810
r: 2270
r: 1459
r: 1462
r: 1505
r: 1508
r: 1677
r: 1680
r: 1886
r: 1314
r: 1318
r: 1339
r: 1485
r: 1488
r: 1596
r: 1601
r: 1606
r: 1611
r: 1636
r: 1641
r: 1760
r: 1782
r: 1882
s: 1319
s: 1452
s: 1462
s: 1472
s: 1489
s: 1498
s: 1508
s: 1525
s: 1534
s: 1551
s: 1560
s: 1600
s: 1610
s: 1630
s: 1640
s: 1650
s: 1661
s: 1671
s: 1680
s: 1702
s: 1546
s: 1557
s: 1721
s: 1469
s: 1725
s: 1505
s: 1597
s: 1881
s: 1607
s: 1886
s: 1627
s: 1637
s: 1761
s: 1647
s: 1656
s: 1666
s: 1677
s: 1769
s: 1781
s: 1810
s: 1315
s: 1459
s: 1484
s: 1495
s: 2269
d: [2,4253,3,0,0,3,52,1886,1702,2269,1300]
e: 4253
d: [2,4256,3,4,2,1,40,2269,1300]
d: [2,4258,3,2,1,2,21,4256,1702]
d: [2,4259,3,0,0,3,21,4258,1886]
e: 4259
r: 1546
r: 1551
r: 1557
r: 1560
r: 1656
r: 1661
r: 1666
r: 1671
r: 1702
r: 1721
r: 1326
r: 1330
r: 1334
r: 1342
r: 1350
r: 1393
r: 1396
r: 1567
r: 1570
r: 1577
r: 1580
r: 1597
r: 1600
r: 1607
r: 1610
r: 1627
r: 1630
r: 1637
r: 1640
r: 1725
r: 1730
r: 1786
r: 1791
r: 1796
r: 1877
r: 2269
s: 1327
s: 1335
s: 1343
s: 1351
s: 1397
s: 1550
s: 1561
s: 1571
s: 1581
s: 1601
s: 1611
s: 1631
s: 1641
s: 1660
s: 1670
s: 1547
s: 1556
s: 1703
s: 1566
s: 1720
s: 1576
s: 1596
s: 1726
s: 1606
s: 1731
s: 1626
s: 1636
s: 1878
s: 1657
s: 1667
s: 1787
s: 1792
s: 1797
s: 1449
s: 2270
s: 1323
s: 1387
s: 1392
d: [2,4341,1,4,2,1,44,1392,1266]
d: [2,4345,1,2,1,2,21,4341,1687]
d: [2,4346,1,0,0,3,21,4345,1670]
e: 4346
r: 1315
r: 1319
r: 1343
r: 1392
r: 1397
r: 1438
r: 1443
r: 1484
r: 1489
r: 1547
r: 1550
r: 1556
r: 1561
r: 1596
r: 1601
r: 1703
r: 1731
r: 1756
r: 1781
r: 1797
r: 2029
r: 1327
r: 1566
r: 1571
r: 1626
r: 1631
r: 1787
s: 1318
s: 1326
s: 1343
s: 1396
s: 1442
s: 1488
s: 1551
s: 1560
s: 1570
s: 1600
s: 1630
s: 1702
s: 1546
s: 1557
s: 1730
s: 1567
s: 1755
s: 1597
s: 1627
a: [3,2028,1,3,1,2025]
s: 2028
d: [2,4396,2,0,0,2,52,1687,2028,1239]
e: 4396
d: [2,4397,2,2,1,1,40,2028,1239]
d: [2,4399,2,0,0,2,21,4397,1687]
e: 4399
r: 1318
r: 1323
r: 1346
r: 1351
r: 1387
r: 1449
r: 1452
r: 1459
r: 1462
r: 1495
r: 1498
r: 1505
r: 1508
r: 1525
r: 1534
r: 1546
r: 1551
r: 1557
r: 1560
r: 1576
r: 1581
r: 1617
r: 1620
r: 1627
r: 1630
r: 1667
r: 1670
r: 1677
r: 1680
r: 1702
r: 1720
r: 1755
r: 1769
r: 1810
r: 1881
r: 1886
r: 2028
r: 2270
s: 1319
s: 1350
s: 1453
s: 1463
s: 1499
s: 1509
s: 1524
s: 1535
s: 1550
s: 1561
s: 1580
s: 1621
s: 1631
s: 1671
s: 1681
s: 1547
s: 1703
s: 1556
s: 1504
s: 1721
s: 1577
s: 1882
s: 1616
s: 1626
s: 1887
s: 1666
s: 1756
s: 1676
s: 1768
s: 2029
s: 1781
s: 1786
s: 1797
s: 1811
s: 1315
s: 1458
s: 1485
s: 1494
s: 1439
s: 2269
s: 1448
s: 1322
s: 1386
s: 1393
s: 2334
s: 1331
s: 2380
s: 1339
s: 1347
d: [2,4500,3,0,0,3,52,1730,1756,1347,2218]
e: 4500
d: [2,4503,2,4,2,1,40,1347,2218]
d: [2,4505,2,2,1,2,21,4503,1756]
d: [2,4506,2,0,0,3,21,4505,1730]
e: 4506
r: 1343
r: 1350
r: 1386
r: 1439
r: 1442
r: 1448
r: 1453
r: 1458
r: 1463
r: 1577
r: 1580
r: 1587
r: 1590
r: 1606
r: 1611
r: 1616
r: 1621
r: 1730
r: 1756
r: 1797
r: 1878
r: 1882
r: 2380
r: 1315
r: 1319
r: 1322
r: 1326
r: 1331
r: 1335
r: 1339
r: 1403
r: 1406
r: 1413
r: 1416
r: 1469
r: 1472
r: 1494
r: 1499
r: 1535
r: 1547
r: 1550
r: 1556
r: 1561
r: 1567
r: 1570
r: 1597
r: 1600
r: 1626
r: 1631
r: 1647
r: 1650
r: 1703
r: 1761
r: 1768
r: 1781
r: 1786
r: 1792
s: 1318
s: 1327
s: 1334
s: 1342
s: 1350
s: 1407
s: 1417
s: 1442
s: 1452
s: 1463
s: 1473
s: 1498
s: 1534
s: 1551
s: 1560
s: 1571
s: 1581
s: 1590
s: 1601
s: 1611
s: 1620
s: 1630
s: 1651
s: 1546
s: 1557
s: 1702
s: 1468
s: 1730
s: 1566
s: 1576
s: 1878
s: 1587
s: 1881
s: 1596
s: 1606
s: 1755
s: 1617
s: 1760
s: 1627
s: 1646
s: 1769
s: 1782
s: 1787
s: 1791
s: 1796
s: 1314
s: 1412
s: 1458
s: 1495
s: 1402
s: 1439
s: 1449
s: 1386
d: [2,4613,1,4,2,1,44,1386,1272]
d: [2,4617,1,2,1,2,21,4613,1687]
d: [2,4618,1,0,0,3,21,4617,1407]
e: 4618
d: [2,4614,1,4,2,1,44,1386,1271]
d: [2,4620,1,2,1,2,21,4614,1687]
d: [2,4621,1,0,0,3,21,4620,1417]
e: 4621
r: 1327
r: 1393
r: 1396
r: 1402
r: 1407
r: 1412
r: 1417
r: 1449
r: 1452
r: 1458
r: 1463
r: 1468
r: 1473
r: 1495
r: 1498
r: 1504
r: 1509
r: 1534
r: 1546
r: 1551
r: 1557
r: 1560
r: 1627
r: 1630
r: 1646
r: 1651
r: 1657
r: 1660
r: 1676
r: 1681
r: 1702
r: 1721
r: 1760
r: 1769
r: 1787
i: [1,1233,0,53,10,1,0,0] t42_gate_1
i: [1,1224,0,53,10,1,0,0] t42_gate_1
i: [1,1221,0,53,10,1,0,0] t42_gate_1
i: [1,1188,0,53,10,1,0,0] t42_gate_1
i: [1,1185,0,53,10,1,0,0] t42_gate_1
i: [1,1176,0,53,10,1,0,0] t42_gate_1
i: [1,1173,0,53,10,1,0,0] t42_gate_1
i: [1,1044,0,53,10,1,0,0] t42_gate_1
i: [1,1041,0,53,10,1,0,0] t42_gate_1
i: [1,1032,0,53,10,1,0,0] t42_gate_1
i: [1,1029,0,53,10,1,0,0] t42_gate_1
i: [1,996,0,53,10,1,0,0] t42_gate_1
i: [1,993,0,53,10,1,0,0] t42_gate_1
i: [1,984,0,53,10,1,0,0] t42_gate_1
i: [1,981,0,53,10,1,0,0] t42_gate_1
i: [1,468,0,53,10,1,0,0] t42_gate_1
i: [1,465,0,53,10,1,0,0] t42_gate_1
i: [1,456,0,53,10,1,0,0] t42_gate_1
i: [1,453,0,53,10,1,0,0] t42_gate_1
i: [1,420,0,53,10,1,0,0] t42_gate_1
i: [1,417,0,53,10,1,0,0] t42_gate_1
i: [1,408,0,53,10,1,0,0] t42_gate_1
i: [1,405,0,53,10,1,0,0] t42_gate_1
i: [1,276,0,53,10,1,0,0] t42_gate_1
i: [1,273,0,53,10,1,0,0] t42_gate_1
i: [1,264,0,53,10,1,0,0] t42_gate_1
i: [1,261,0,53,10,1,0,0] t42_gate_1
i: [1,228,0,53,10,1,0,0] t42_gate_1
i: [1,225,0,53,10,1,0,0] t42_gate_1
i: [1,216,0,53,10,1,0,0] t42_gate_1
i: [1,213,0,53,10,1,0,0] t42_gate_1
i: [1,211,0,9,2,1,0,0] t42_gate_1
i: [1,209,0,9,2,1,0,0] t42_gate_1
i: [1,207,0,8,2,1,0,0] t42_gate_1
i: [1,205,0,7,2,1,0,0] t42_gate_1
i: [1,203,0,6,2,1,0,0] t42_gate_1
i: [1,201,0,6,2,1,0,0] t42_gate_1
i: [1,199,0,8,2,1,0,0] t42_gate_1
i: [1,197,0,8,2,1,0,0] t42_gate_1
i: [1,195,0,7,2,1,0,0] t42_gate_1
i: [1,193,0,6,2,1,0,0] t42_gate_1
i: [1,191,0,6,2,1,0,0] t42_gate_1
i: [1,189,0,7,2,1,0,0] t42_gate_1
i: [1,187,0,7,2,1,0,0] t42_gate_1
i: [1,185,0,6,2,1,0,0] t42_gate_1
i: [1,183,0,6,2,1,0,0] t42_gate_1
i: [1,181,0,6,2,1,0,0] t42_gate_1
i: [1,179,0,6,2,1,0,0] t42_gate_1
i: [1,177,0,6,2,1,0,0] t42_gate_1
i: [1,175,0,5,2,1,0,0] t42_gate_1
i: [1,173,0,5,2,1,0,0] t42_gate_1
i: [1,171,0,6,2,1,0,0] t42_gate_1
i: [1,169,0,6,2,1,0,0] t42_gate_1
i: [1,167,0,6,2,1,0,0] t42_gate_1
i: [1,165,0,6,2,1,0,0] t42_gate_1
i: [1,163,0,6,2,1,0,0] t42_gate_1
i: [1,161,0,6,2,1,0,0] t42_gate_1
i: [1,159,0,6,2,1,0,0] t42_gate_1
i: [1,157,0,6,2,1,0,0] t42_gate_1
i: [1,155,0,6,2,1,0,0] t42_gate_1
i: [1,153,0,6,2,1,0,0] t42_gate_1
i: [1,151,0,6,2,1,0,0] t42_gate_1
i: [1,149,0,6,2,1,0,0] t42_gate_1
d: [2,1733,1,2,1,1,40,1284,1577]
d: [2,1734,1,2,1,1,40,1284,1637]
d: [2,1717,1,2,1,1,40,1283,1520]
d: [2,1850,2,2,1,1,40,1468,1764]
d: [2,1824,1,2,1,1,40,1293,1412]
d: [2,1893,1,2,1,1,40,1607,1284]
d: [2,1817,1,2,1,1,40,1810,1283]
d: [2,2413,3,9,2,1,40,2379,2182]
d: [2,1998,1,4,2,1,44,1438,1259]
d: [2,2592,1,2,1,2,21,1998,1687]
d: [2,2119,1,4,2,1,40,1505,1286]
d: [2,2707,1,2,1,2,21,2119,1641]
d: [2,2634,2,11,3,1,40,1330,2184]
d: [2,2709,2,4,2,2,21,2634,2333]
d: [2,2710,2,2,1,3,21,2709,1881]
d: [2,1713,1,2,1,1,40,1546,1238]
d: [2,2546,1,6,3,1,40,1315,1303]
d: [2,2274,3,4,2,1,40,2270,1307]
d: [2,2508,1,6,3,1,44,1656,1244]
d: [2,2741,1,4,2,2,21,2508,1687]
d: [2,2507,1,6,3,1,44,1656,1245]
d: [2,2522,1,4,2,2,21,2507,1687]
d: [2,2544,1,6,3,1,40,1315,1305]
d: [2,2543,1,4,2,1,40,1666,1287]
d: [2,2642,3,7,2,1,40,2333,2145]
d: [2,2764,3,5,1,2,21,2642,1760]
d: [2,2414,3,9,2,1,40,2379,2181]
d: [2,2768,3,7,1,2,21,2414,1878]
d: [2,2895,2,7,2,2,21,2889,1761]
d: [2,2896,2,5,1,3,21,2895,1725]
d: [2,2957,1,4,2,1,40,1797,1289]
d: [2,2956,2,4,2,1,40,1797,1927]
d: [2,2902,1,4,2,1,44,1386,1270]
d: [2,3042,1,2,1,2,21,2902,1687]
d: [2,3089,3,2,1,2,21,3070,1725]
d: [2,2833,1,4,2,1,44,1484,1248]
d: [2,3093,1,2,1,2,21,2833,1687]
d: [2,3212,2,9,3,1,40,1323,2151]
d: [2,3214,2,7,2,2,21,3212,1760]
d: [2,3215,2,5,1,3,21,3214,1726]
d: [2,2640,2,13,3,1,40,2225,1338]
d: [2,2638,2,13,3,1,40,1338,2224]
d: [2,3069,3,7,2,1,40,2334,2140]
d: [2,3421,3,5,1,2,21,3415,1726]
d: [2,3319,2,9,3,1,40,1322,2148]
d: [2,3423,2,7,2,2,21,3319,1726]
d: [2,3424,2,5,1,3,21,3423,1761]
d: [2,3461,3,4,2,1,40,2270,1309]
d: [2,3416,3,4,2,1,40,2333,2143]
d: [2,3509,1,4,2,1,40,1647,1290]
d: [2,3501,1,4,2,1,40,1469,1286]
d: [2,3462,3,4,2,1,40,2270,1308]
d: [2,3363,2,11,2,2,21,3352,1755]
d: [2,3022,2,11,3,1,40,1331,2187]
d: [2,3023,2,11,3,1,40,1331,2183]
d: [2,2788,1,4,2,1,40,1677,1290]
d: [2,3690,1,6,3,1,44,1626,1244]
d: [2,3691,1,4,2,2,21,3690,1687]
d: [2,3692,1,2,1,3,21,3691,1726]
d: [2,3716,1,2,1,1,40,1494,1294]
d: [2,3826,1,4,2,1,44,1484,1247]
d: [2,3839,1,2,1,2,21,3826,1687]
d: [2,3822,1,6,3,1,40,1459,1295]
d: [2,3714,1,6,3,1,40,1314,1297]
d: [2,3852,1,4,2,2,21,3714,1702]
d: [2,3713,1,6,3,1,40,1314,1298]
d: [2,3892,2,2,1,1,40,1607,1747]
d: [2,3963,2,13,3,1,40,1339,2229]
d: [2,3978,2,11,2,2,21,3963,1730]
d: [2,4006,1,4,2,1,40,1787,1289]
d: [2,3942,1,11,5,1,44,1449,1250]
d: [2,3972,1,8,4,2,21,3942,2029]
d: [2,4030,1,6,3,3,21,3972,1611]
d: [2,4005,2,4,2,1,40,1787,1927]
d: [2,3997,1,6,3,1,44,1596,1244]
d: [2,4037,1,4,2,2,21,3997,1687]
d: [2,3895,1,4,2,1,40,1617,1290]
d: [2,4048,1,6,3,1,44,1626,1245]
d: [2,4050,1,4,2,2,21,4048,1687]
d: [2,4051,1,2,1,3,21,4050,1761]
d: [2,4053,1,4,2,1,40,1636,1287]
d: [2,4054,1,2,1,2,21,4053,1761]
d: [2,4066,2,9,3,1,40,1323,2152]
d: [2,4069,2,7,2,2,21,4066,1761]
d: [2,4070,2,5,1,3,21,4069,1726]
d: [2,4087,2,11,3,1,40,1330,2190]
d: [2,4095,2,9,2,2,21,4087,1882]
d: [2,4112,1,4,2,1,40,1796,1288]
d: [2,4109,1,4,2,1,40,1791,1291]
d: [2,4111,1,4,2,1,40,1796,1291]
d: [2,4110,1,4,2,1,40,1791,1288]
d: [2,3715,1,6,3,1,40,1413,1295]
d: [2,4144,1,4,2,2,21,3715,1611]
d: [2,4124,2,11,2,1,40,1346,2222]
d: [2,4177,1,4,2,1,44,1438,1258]
d: [2,4183,1,2,1,2,21,4177,1687]
d: [2,4169,1,4,2,1,40,1786,1291]
d: [2,4170,1,4,2,1,40,1786,1288]
d: [2,4178,1,4,2,1,44,1438,1257]
d: [2,4195,1,2,1,2,21,4178,1687]
d: [2,4089,2,11,3,1,40,1330,2185]
d: [2,4200,2,9,2,2,21,4089,1877]
d: [2,4088,2,11,3,1,40,1330,2186]
d: [2,4202,2,9,2,2,21,4088,1877]
d: [2,4206,0,4,1,1,22,1520,1282]
d: [2,4221,1,7,3,1,44,1627,1242]
d: [2,4225,1,4,2,2,21,4221,2029]
d: [2,4226,1,2,1,3,21,4225,1725]
d: [2,4228,2,2,1,1,40,1637,1747]
d: [2,4239,2,4,2,1,40,1781,1927]
d: [2,4238,1,4,2,1,40,1769,1286]
d: [2,4240,1,4,2,1,40,1781,1289]
d: [2,4264,3,2,1,2,21,4256,1886]
d: [2,4244,1,6,3,1,40,1315,1304]
d: [2,4266,1,4,2,2,21,4244,1886]
d: [2,4255,3,4,2,1,40,2269,1301]
d: [2,4268,3,2,1,2,21,4255,1886]
d: [2,4246,1,6,3,1,40,1315,1296]
d: [2,4271,1,4,2,2,21,4246,1886]
d: [2,4252,1,6,3,1,40,1495,1295]
d: [2,4125,2,11,2,1,40,1346,2221]
d: [2,4277,0,4,1,1,22,1479,1281]
d: [2,4282,1,2,1,1,40,1556,1238]
d: [2,4296,1,6,3,1,44,1596,1245]
d: [2,4298,1,4,2,2,21,4296,1687]
d: [2,4299,1,2,1,3,21,4298,1881]
d: [2,4301,1,4,2,1,40,1606,1287]
d: [2,4302,1,2,1,2,21,4301,1881]
d: [2,4333,1,13,6,1,44,1387,1268]
d: [2,4348,1,11,5,2,21,4333,1406]
d: [2,4349,1,9,4,3,21,4348,1416]
d: [2,4350,1,7,3,4,21,4349,1426]
d: [2,4351,1,5,2,5,21,4350,1590]
d: [2,4317,2,4,2,1,40,1792,1927]
d: [2,4318,1,4,2,1,40,1792,1289]
d: [2,4287,1,6,3,1,44,1566,1244]
d: [2,4356,1,4,2,2,21,4287,1687]
d: [2,3676,1,4,2,1,40,1587,1290]
d: [2,4126,2,4,2,1,40,1346,2220]
d: [2,4372,0,4,1,1,22,1433,1281]
d: [2,4390,1,7,3,1,44,1627,1241]
d: [2,4393,1,5,2,2,21,4390,1726]
d: [2,4394,1,3,1,3,21,4393,1761]
d: [2,4149,1,4,2,1,40,1423,1286]
d: [2,4310,1,7,3,1,44,1657,1241]
d: [2,4309,1,7,3,1,44,1657,1242]
a: [3,1541,0,4,1,178]
d: [2,4408,0,4,1,1,22,1541,1281]
d: [2,4412,0,4,1,1,22,1521,1282]
d: [2,4418,2,2,1,1,40,1577,1747]
d: [2,3929,1,11,5,1,44,1403,1250]
d: [2,4402,1,9,4,2,21,3929,1641]
d: [2,4513,1,6,3,3,21,4402,2029]
d: [2,4514,1,4,2,4,21,4513,1681]
d: [2,4497,2,13,3,1,40,1339,2228]
d: [2,4377,1,7,3,1,44,1567,1241]
d: [2,4517,1,4,2,2,21,4377,2029]
d: [2,4376,1,7,3,1,44,1567,1242]
d: [2,4519,1,4,2,2,21,4376,2029]
d: [2,4489,2,11,3,1,40,1331,2189]
d: [2,4523,2,4,2,2,21,4489,2334]
d: [2,4382,1,7,3,1,44,1597,1242]
d: [2,4525,1,4,2,2,21,4382,2029]
d: [2,4383,1,7,3,1,44,1597,1241]
d: [2,4527,1,4,2,2,21,4383,2029]
d: [2,4490,2,11,3,1,40,1331,2188]
d: [2,4529,2,4,2,2,21,4490,2334]
d: [2,4502,2,11,2,1,40,1347,2217]
d: [2,4501,2,11,2,1,40,1347,2219]
d: [2,4499,2,13,3,1,40,1339,2223]
d: [2,4498,2,13,3,1,40,1339,2227]
a: [3,1432,0,4,1,202]
d: [2,4535,0,4,1,1,22,1432,1281]
d: [2,4480,1,13,6,1,44,1393,1262]
d: [2,4508,1,10,5,2,21,4480,2029]
d: [2,4537,1,8,4,3,21,4508,1671]
d: [2,4538,1,6,3,4,21,4537,1641]
d: [2,4540,0,4,1,1,22,1540,1281]
d: [2,4542,0,4,1,1,22,1531,1280]
d: [2,4544,0,4,1,1,22,1515,1281]
a: [3,1478,0,4,1,192]
d: [2,4546,0,4,1,1,22,1478,1281]
d: [2,4549,1,2,1,1,40,1557,1239]
d: [2,4555,1,6,3,1,44,1566,1245]
d: [2,4557,1,4,2,2,21,4555,1687]
d: [2,4558,1,2,1,3,21,4557,1730]
d: [2,4560,1,4,2,1,40,1576,1287]
d: [2,4561,1,2,1,2,21,4560,1730]
d: [2,4612,1,4,2,1,44,1386,1273]
d: [2,4623,1,2,1,2,21,4612,1687]
d: [2,4601,1,11,5,1,44,1439,1255]
d: [2,4627,1,8,4,2,21,4601,2029]
d: [2,4628,1,6,3,3,21,4627,1620]
d: [2,4450,1,9,4,1,44,1485,1246]
d: [2,4630,1,6,3,2,21,4450,2029]
d: [2,4577,1,4,2,1,40,1782,1291]
d: [2,4578,1,4,2,1,40,1782,1288]
d: [2,4442,1,4,2,1,40,1811,1290]
d: [2,4585,1,6,3,1,40,1314,1306]
d: [2,4635,1,4,2,2,21,4585,1887]
d: [2,4586,1,6,3,1,40,1314,1299]
d: [2,4637,1,4,2,2,21,4586,1887]
d: [2,4462,3,4,2,1,40,2269,1302]
a: [3,1530,0,4,1,180]
d: [2,4640,0,4,1,1,22,1530,1280]
a: [3,1514,0,4,1,184]
d: [2,4642,0,4,1,1,22,1514,1281]
f: 4642,4640,4462,4637,4635,4442,4578,4577,4630,4628,4620,4617,4623,4561,4558,4549,4546,4544,4542,4540,4538,4535,4498,4499,4501,4502,4529,4527,4525,4523,4503,4519,4517,4497,4514,4418,4412,4408,4309,4310,4149,4399,4394,4372,4126,3676,4356,4318,4317,4351,4346,4302,4299,4282,4277,4125,4252,4271,4268,4266,4264,4240,4238,4239,4228,4226,4206,4202,4200,4195,4170,4169,4180,4183,4124,4144,4110,4111,4109,4112,4138,4095,4070,4054,4051,3895,4037,4005,4030,4006,4027,3978,3892,3714,3713,3866,3852,3822,3839,3837,3834,3816,3774,3716,3692,2788,3023,3022,3363,3462,3501,3625,3509,3589,3569,3520,3416,3461,3480,3424,3421,3069,3240,2638,2647,2640,3215,3167,3093,3089,3042,3012,3004,2956,2957,2977,2934,2896,2865,2862,2793,2783,2768,2764,2543,2544,2522,2741,2274,2546,1713,2732,2710,2707,2694,2696,2688,2685,2616,2592,2587,2581,2578,2566,2541,2413,2465,2400,2397,2394,2266,2126,1817,2054,1893,1824,1850,1854,1826,1819,1717,1745,1743,1734,1733,1716,1715,1714,1710,1237,149,150,151,152,153,154,155,156,157,158,159,160,161,162,163,164,165,166,167,168,169,170,171,172,173,174,175,176,177,178,179,180,181,182,183,184,185,186,187,188,189,190,191,192,193,194,195,196,197,198,199,200,201,202,203,204,205,206,207,208,209,210,211,212,213,216,225,228,261,264,273,276,405,408,417,420,453,456,465,468,981,984,993,996,1029,1032,1041,1044,1173,1176,1185,1188,1221,1224,1233,1236
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t42_gate_1
% SZS output start Proof for gate_1__t42_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ! [X22] : ! [X23] : ! [X24] : ! [X25] : ! [X26] : ! [X27] : ! [X28] : ! [X29] : ! [X30] : ! [X31] : ! [X32] : ! [X33] : ! [X34] : ! [X35] : ! [X36] : ! [X37] : ! [X38] : ! [X39] : ! [X40] : ~(~(~(v1_xboole_0(X9) & ~v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) & ~(v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X40) & ~v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) & ~(v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X40)) & ~(v1_xboole_0(X39) & ~v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0))) & ~(v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0)) & ~v1_xboole_0(X39)) & ~(v1_xboole_0(X38) & ~v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0))) & ~(v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0)) & ~v1_xboole_0(X38)) & ~(v1_xboole_0(X37) & ~v1_xboole_0(k10_gate_1(X1,X2,X0))) & ~(v1_xboole_0(k10_gate_1(X1,X2,X0)) & ~v1_xboole_0(X37))) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36))) & ~(v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X36) & ~v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22))) & ~(v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22)) & ~v1_xboole_0(X36)) & ~(v1_xboole_0(X35) & ~v1_xboole_0(k14_gate_1(X10,X20,X17,X14))) & ~(v1_xboole_0(k14_gate_1(X10,X20,X17,X14)) & ~v1_xboole_0(X35)) & ~(v1_xboole_0(X34) & ~v1_xboole_0(k8_gate_1(X13,X20,X17))) & ~(v1_xboole_0(k8_gate_1(X13,X20,X17)) & ~v1_xboole_0(X34)) & ~(v1_xboole_0(X33) & ~v1_xboole_0(k2_gate_1(X16,X20))) & ~(v1_xboole_0(k2_gate_1(X16,X20)) & ~v1_xboole_0(X33)) & ~(v1_xboole_0(X40) & ~v1_xboole_0(k4_gate_1(X32,X21))) & ~(v1_xboole_0(k4_gate_1(X32,X21)) & ~v1_xboole_0(X40)) & ~(v1_xboole_0(X32) & ~v1_xboole_0(k17_gate_1(X29,X30,X31,X16))) & ~(v1_xboole_0(k17_gate_1(X29,X30,X31,X16)) & ~v1_xboole_0(X32)) & ~(v1_xboole_0(X31) & ~v1_xboole_0(k14_gate_1(X17,X14,X11,X22))) & ~(v1_xboole_0(k14_gate_1(X17,X14,X11,X22)) & ~v1_xboole_0(X31)) & ~(v1_xboole_0(X30) & ~v1_xboole_0(k8_gate_1(X10,X17,X14))) & ~(v1_xboole_0(k8_gate_1(X10,X17,X14)) & ~v1_xboole_0(X30)) & ~(v1_xboole_0(X29) & ~v1_xboole_0(k2_gate_1(X13,X17))) & ~(v1_xboole_0(k2_gate_1(X13,X17)) & ~v1_xboole_0(X29)) & ~(v1_xboole_0(X39) & ~v1_xboole_0(k4_gate_1(X28,X18))) & ~(v1_xboole_0(k4_gate_1(X28,X18)) & ~v1_xboole_0(X39)) & ~(v1_xboole_0(X28) & ~v1_xboole_0(k13_gate_1(X26,X27,X13))) & ~(v1_xboole_0(k13_gate_1(X26,X27,X13)) & ~v1_xboole_0(X28)) & ~(v1_xboole_0(X27) & ~v1_xboole_0(k8_gate_1(X14,X11,X22))) & ~(v1_xboole_0(k8_gate_1(X14,X11,X22)) & ~v1_xboole_0(X27)) & ~(v1_xboole_0(X26) & ~v1_xboole_0(k2_gate_1(X10,X14))) & ~(v1_xboole_0(k2_gate_1(X10,X14)) & ~v1_xboole_0(X26)) & ~(v1_xboole_0(X38) & ~v1_xboole_0(k4_gate_1(X25,X15))) & ~(v1_xboole_0(k4_gate_1(X25,X15)) & ~v1_xboole_0(X38)) & ~(v1_xboole_0(X25) & ~v1_xboole_0(k7_gate_1(X24,X10))) & ~(v1_xboole_0(k7_gate_1(X24,X10)) & ~v1_xboole_0(X25)) & ~(v1_xboole_0(X24) & ~v1_xboole_0(k2_gate_1(X22,X11))) & ~(v1_xboole_0(k2_gate_1(X22,X11)) & ~v1_xboole_0(X24)) & ~(v1_xboole_0(X37) & ~v1_xboole_0(k4_gate_1(X23,X12))) & ~(v1_xboole_0(k4_gate_1(X23,X12)) & ~v1_xboole_0(X37)) & ~(v1_xboole_0(X23) & ~v1_xboole_0(k1_gate_1(X22))) & ~(v1_xboole_0(k1_gate_1(X22)) & ~v1_xboole_0(X23)) & ~(v1_xboole_0(X22) & ~v1_xboole_0(k1_gate_1(X0))) & ~(v1_xboole_0(k1_gate_1(X0)) & ~v1_xboole_0(X22)) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k34_gate_1(X7,X8))) & ~(v1_xboole_0(k34_gate_1(X7,X8)) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k6_gate_1(X7,X8))) & ~(v1_xboole_0(k6_gate_1(X7,X8)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k7_gate_1(X7,X8))) & ~(v1_xboole_0(k7_gate_1(X7,X8)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X18) & ~v1_xboole_0(k34_gate_1(X5,X6))) & ~(v1_xboole_0(k34_gate_1(X5,X6)) & ~v1_xboole_0(X18)) & ~(v1_xboole_0(X17) & ~v1_xboole_0(k6_gate_1(X5,X6))) & ~(v1_xboole_0(k6_gate_1(X5,X6)) & ~v1_xboole_0(X17)) & ~(v1_xboole_0(X16) & ~v1_xboole_0(k7_gate_1(X5,X6))) & ~(v1_xboole_0(k7_gate_1(X5,X6)) & ~v1_xboole_0(X16)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k34_gate_1(X3,X4))) & ~(v1_xboole_0(k34_gate_1(X3,X4)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k6_gate_1(X3,X4))) & ~(v1_xboole_0(k6_gate_1(X3,X4)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k7_gate_1(X3,X4))) & ~(v1_xboole_0(k7_gate_1(X3,X4)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k34_gate_1(X1,X2))) & ~(v1_xboole_0(k34_gate_1(X1,X2)) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k6_gate_1(X1,X2))) & ~(v1_xboole_0(k6_gate_1(X1,X2)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k7_gate_1(X1,X2))) & ~(v1_xboole_0(k7_gate_1(X1,X2)) & ~v1_xboole_0(X10))) [input t42_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ! [X22] : ! [X23] : ! [X24] : ! [X25] : ! [X26] : ! [X27] : ! [X28] : ! [X29] : ! [X30] : ! [X31] : ! [X32] : ! [X33] : ! [X34] : ! [X35] : ! [X36] : ! [X37] : ! [X38] : ! [X39] : ! [X40] : ~(~(~(v1_xboole_0(X9) & ~v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) & ~(v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X40) & ~v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) & ~(v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X40)) & ~(v1_xboole_0(X39) & ~v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0))) & ~(v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0)) & ~v1_xboole_0(X39)) & ~(v1_xboole_0(X38) & ~v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0))) & ~(v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0)) & ~v1_xboole_0(X38)) & ~(v1_xboole_0(X37) & ~v1_xboole_0(k10_gate_1(X1,X2,X0))) & ~(v1_xboole_0(k10_gate_1(X1,X2,X0)) & ~v1_xboole_0(X37))) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36))) & ~(v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X36) & ~v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22))) & ~(v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22)) & ~v1_xboole_0(X36)) & ~(v1_xboole_0(X35) & ~v1_xboole_0(k14_gate_1(X10,X20,X17,X14))) & ~(v1_xboole_0(k14_gate_1(X10,X20,X17,X14)) & ~v1_xboole_0(X35)) & ~(v1_xboole_0(X34) & ~v1_xboole_0(k8_gate_1(X13,X20,X17))) & ~(v1_xboole_0(k8_gate_1(X13,X20,X17)) & ~v1_xboole_0(X34)) & ~(v1_xboole_0(X33) & ~v1_xboole_0(k2_gate_1(X16,X20))) & ~(v1_xboole_0(k2_gate_1(X16,X20)) & ~v1_xboole_0(X33)) & ~(v1_xboole_0(X40) & ~v1_xboole_0(k4_gate_1(X32,X21))) & ~(v1_xboole_0(k4_gate_1(X32,X21)) & ~v1_xboole_0(X40)) & ~(v1_xboole_0(X32) & ~v1_xboole_0(k17_gate_1(X29,X30,X31,X16))) & ~(v1_xboole_0(k17_gate_1(X29,X30,X31,X16)) & ~v1_xboole_0(X32)) & ~(v1_xboole_0(X31) & ~v1_xboole_0(k14_gate_1(X17,X14,X11,X22))) & ~(v1_xboole_0(k14_gate_1(X17,X14,X11,X22)) & ~v1_xboole_0(X31)) & ~(v1_xboole_0(X30) & ~v1_xboole_0(k8_gate_1(X10,X17,X14))) & ~(v1_xboole_0(k8_gate_1(X10,X17,X14)) & ~v1_xboole_0(X30)) & ~(v1_xboole_0(X29) & ~v1_xboole_0(k2_gate_1(X13,X17))) & ~(v1_xboole_0(k2_gate_1(X13,X17)) & ~v1_xboole_0(X29)) & ~(v1_xboole_0(X39) & ~v1_xboole_0(k4_gate_1(X28,X18))) & ~(v1_xboole_0(k4_gate_1(X28,X18)) & ~v1_xboole_0(X39)) & ~(v1_xboole_0(X28) & ~v1_xboole_0(k13_gate_1(X26,X27,X13))) & ~(v1_xboole_0(k13_gate_1(X26,X27,X13)) & ~v1_xboole_0(X28)) & ~(v1_xboole_0(X27) & ~v1_xboole_0(k8_gate_1(X14,X11,X22))) & ~(v1_xboole_0(k8_gate_1(X14,X11,X22)) & ~v1_xboole_0(X27)) & ~(v1_xboole_0(X26) & ~v1_xboole_0(k2_gate_1(X10,X14))) & ~(v1_xboole_0(k2_gate_1(X10,X14)) & ~v1_xboole_0(X26)) & ~(v1_xboole_0(X38) & ~v1_xboole_0(k4_gate_1(X25,X15))) & ~(v1_xboole_0(k4_gate_1(X25,X15)) & ~v1_xboole_0(X38)) & ~(v1_xboole_0(X25) & ~v1_xboole_0(k7_gate_1(X24,X10))) & ~(v1_xboole_0(k7_gate_1(X24,X10)) & ~v1_xboole_0(X25)) & ~(v1_xboole_0(X24) & ~v1_xboole_0(k2_gate_1(X22,X11))) & ~(v1_xboole_0(k2_gate_1(X22,X11)) & ~v1_xboole_0(X24)) & ~(v1_xboole_0(X37) & ~v1_xboole_0(k4_gate_1(X23,X12))) & ~(v1_xboole_0(k4_gate_1(X23,X12)) & ~v1_xboole_0(X37)) & ~(v1_xboole_0(X23) & ~v1_xboole_0(k1_gate_1(X22))) & ~(v1_xboole_0(k1_gate_1(X22)) & ~v1_xboole_0(X23)) & ~(v1_xboole_0(X22) & ~v1_xboole_0(k1_gate_1(X0))) & ~(v1_xboole_0(k1_gate_1(X0)) & ~v1_xboole_0(X22)) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k34_gate_1(X7,X8))) & ~(v1_xboole_0(k34_gate_1(X7,X8)) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k6_gate_1(X7,X8))) & ~(v1_xboole_0(k6_gate_1(X7,X8)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k7_gate_1(X7,X8))) & ~(v1_xboole_0(k7_gate_1(X7,X8)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X18) & ~v1_xboole_0(k34_gate_1(X5,X6))) & ~(v1_xboole_0(k34_gate_1(X5,X6)) & ~v1_xboole_0(X18)) & ~(v1_xboole_0(X17) & ~v1_xboole_0(k6_gate_1(X5,X6))) & ~(v1_xboole_0(k6_gate_1(X5,X6)) & ~v1_xboole_0(X17)) & ~(v1_xboole_0(X16) & ~v1_xboole_0(k7_gate_1(X5,X6))) & ~(v1_xboole_0(k7_gate_1(X5,X6)) & ~v1_xboole_0(X16)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k34_gate_1(X3,X4))) & ~(v1_xboole_0(k34_gate_1(X3,X4)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k6_gate_1(X3,X4))) & ~(v1_xboole_0(k6_gate_1(X3,X4)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k7_gate_1(X3,X4))) & ~(v1_xboole_0(k7_gate_1(X3,X4)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k34_gate_1(X1,X2))) & ~(v1_xboole_0(k34_gate_1(X1,X2)) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k6_gate_1(X1,X2))) & ~(v1_xboole_0(k6_gate_1(X1,X2)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k7_gate_1(X1,X2))) & ~(v1_xboole_0(k7_gate_1(X1,X2)) & ~v1_xboole_0(X10))) [negated conjecture 1]
3. ! [X0,X1] : k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [input commutativity_k2_gate_1]
5. ! [X0,X1] : k4_gate_1(X0,X1) = k4_gate_1(X1,X0) [input commutativity_k4_gate_1]
6. ! [X0,X1] : k6_gate_1(X0,X1) = k6_gate_1(X1,X0) [input commutativity_k6_gate_1]
7. ! [X0,X1] : k7_gate_1(X0,X1) = k7_gate_1(X1,X0) [input commutativity_k7_gate_1]
8. ! [X0] : ! [X1] : ! [X2] : ((~(v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k13_gate_1(X0,X1,X2) = k1_xboole_0) & ~(k13_gate_1(X0,X1,X2) != k1_gate_1(k1_xboole_0) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [input d13_gate_1]
9. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((~(~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k1_xboole_0 = k14_gate_1(X0,X1,X2,X3)) & ~(k1_gate_1(k1_xboole_0) != k14_gate_1(X0,X1,X2,X3) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input d14_gate_1]
10. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k1_xboole_0 = k17_gate_1(X0,X1,X2,X3)) & ~(k1_gate_1(k1_xboole_0) != k17_gate_1(X0,X1,X2,X3) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [input d17_gate_1]
11. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((~(~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k1_xboole_0 = k18_gate_1(X0,X1,X2,X3,X4)) & ~(k1_gate_1(k1_xboole_0) != k18_gate_1(X0,X1,X2,X3,X4) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input d18_gate_1]
12. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k1_xboole_0 = k21_gate_1(X0,X1,X2,X3,X4)) & ~(k1_gate_1(k1_xboole_0) != k21_gate_1(X0,X1,X2,X3,X4) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [input d21_gate_1]
13. ! [X0] : ! [X1] : ((~(~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0))) => k34_gate_1(X0,X1) = k1_xboole_0) & ~(k34_gate_1(X0,X1) != k1_gate_1(k1_xboole_0) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0)))) [input d34_gate_1]
14. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : k35_gate_1(X0,X1,X2,X3,X4) = k10_gate_1(X2,X3,k11_gate_1(X0,X1,X4)) [input d35_gate_1]
15. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : k36_gate_1(X0,X1,X2,X3,X4) = k11_gate_1(X2,X3,k11_gate_1(X0,X1,X4)) [input d36_gate_1]
16. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : k37_gate_1(X0,X1,X2,X3,X4,X5,X6) = k10_gate_1(X4,X5,k36_gate_1(X0,X1,X2,X3,X6)) [input d37_gate_1]
17. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : k38_gate_1(X0,X1,X2,X3,X4,X5,X6) = k11_gate_1(X4,X5,k36_gate_1(X0,X1,X2,X3,X6)) [input d38_gate_1]
18. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : k39_gate_1(X0,X1,X2,X3,X4,X5,X6,X7,X8) = k10_gate_1(X6,X7,k38_gate_1(X0,X1,X2,X3,X4,X5,X8)) [input d39_gate_1]
19. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : k40_gate_1(X0,X1,X2,X3,X4,X5,X6,X7,X8) = k11_gate_1(X6,X7,k38_gate_1(X0,X1,X2,X3,X4,X5,X8)) [input d40_gate_1]
41. ! [X0,X1] : (v1_xboole_0(X0) => ~v1_xboole_0(k6_gate_1(X0,X1))) [input fc13_gate_1]
42. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => v1_xboole_0(k6_gate_1(X0,X1))) [input fc14_gate_1]
43. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => ~v1_xboole_0(k7_gate_1(X0,X1))) [input fc15_gate_1]
44. ! [X0,X1] : (~v1_xboole_0(X0) => v1_xboole_0(k7_gate_1(X0,X1))) [input fc16_gate_1]
45. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc17_gate_1]
46. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc18_gate_1]
47. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X0,X2))) [input fc19_gate_1]
48. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
49. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
50. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X2,X0))) [input fc20_gate_1]
51. ! [X0,X1,X2] : ((v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k10_gate_1(X0,X1,X2))) [input fc25_gate_1]
52. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => ~v1_xboole_0(k10_gate_1(X0,X1,X2))) [input fc26_gate_1]
53. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => ~v1_xboole_0(k10_gate_1(X0,X2,X1))) [input fc27_gate_1]
54. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => ~v1_xboole_0(k10_gate_1(X2,X0,X1))) [input fc28_gate_1]
55. ! [X0,X1,X2] : ((v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => v1_xboole_0(k10_gate_1(X0,X1,X2))) [input fc29_gate_1]
56. ! [X0] : (~v1_xboole_0(X0) => v1_xboole_0(k1_gate_1(X0))) [input fc2_gate_1]
57. ! [X0,X1,X2] : ((v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => v1_xboole_0(k10_gate_1(X0,X2,X1))) [input fc30_gate_1]
58. ! [X0,X1,X2] : ((v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => v1_xboole_0(k10_gate_1(X2,X0,X1))) [input fc31_gate_1]
59. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k10_gate_1(X0,X1,X2))) [input fc32_gate_1]
60. ! [X0,X1,X2] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k11_gate_1(X0,X1,X2))) [input fc33_gate_1]
61. ! [X0,X1,X2] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k11_gate_1(X0,X2,X1))) [input fc34_gate_1]
62. ! [X0,X1,X2] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k11_gate_1(X2,X0,X1))) [input fc35_gate_1]
63. ! [X0,X1,X2] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k11_gate_1(X0,X1,X2))) [input fc36_gate_1]
64. ! [X0,X1,X2] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k11_gate_1(X0,X2,X1))) [input fc37_gate_1]
65. ! [X0,X1,X2] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k11_gate_1(X2,X0,X1))) [input fc38_gate_1]
66. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_gate_1(X0,X1))) [input fc3_gate_1]
67. ! [X0,X1] : (v1_xboole_0(X0) => v1_xboole_0(k2_gate_1(X0,X1))) [input fc4_gate_1]
68. ! [X0,X1] : ((~v1_xboole_0(X1) & v1_xboole_0(X0)) => ~v1_xboole_0(k4_gate_1(X0,X1))) [input fc7_gate_1]
69. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k4_gate_1(X0,X1))) [input fc8_gate_1]
70. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => v1_xboole_0(k4_gate_1(X0,X1))) [input fc9_gate_1]
73. ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40] : ~(~(~(v1_xboole_0(X9) & ~v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) & ~(v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X40) & ~v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) & ~(v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X40)) & ~(v1_xboole_0(X39) & ~v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0))) & ~(v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0)) & ~v1_xboole_0(X39)) & ~(v1_xboole_0(X38) & ~v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0))) & ~(v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0)) & ~v1_xboole_0(X38)) & ~(v1_xboole_0(X37) & ~v1_xboole_0(k10_gate_1(X1,X2,X0))) & ~(v1_xboole_0(k10_gate_1(X1,X2,X0)) & ~v1_xboole_0(X37))) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36))) & ~(v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X36) & ~v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22))) & ~(v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22)) & ~v1_xboole_0(X36)) & ~(v1_xboole_0(X35) & ~v1_xboole_0(k14_gate_1(X10,X20,X17,X14))) & ~(v1_xboole_0(k14_gate_1(X10,X20,X17,X14)) & ~v1_xboole_0(X35)) & ~(v1_xboole_0(X34) & ~v1_xboole_0(k8_gate_1(X13,X20,X17))) & ~(v1_xboole_0(k8_gate_1(X13,X20,X17)) & ~v1_xboole_0(X34)) & ~(v1_xboole_0(X33) & ~v1_xboole_0(k2_gate_1(X16,X20))) & ~(v1_xboole_0(k2_gate_1(X16,X20)) & ~v1_xboole_0(X33)) & ~(v1_xboole_0(X40) & ~v1_xboole_0(k4_gate_1(X32,X21))) & ~(v1_xboole_0(k4_gate_1(X32,X21)) & ~v1_xboole_0(X40)) & ~(v1_xboole_0(X32) & ~v1_xboole_0(k17_gate_1(X29,X30,X31,X16))) & ~(v1_xboole_0(k17_gate_1(X29,X30,X31,X16)) & ~v1_xboole_0(X32)) & ~(v1_xboole_0(X31) & ~v1_xboole_0(k14_gate_1(X17,X14,X11,X22))) & ~(v1_xboole_0(k14_gate_1(X17,X14,X11,X22)) & ~v1_xboole_0(X31)) & ~(v1_xboole_0(X30) & ~v1_xboole_0(k8_gate_1(X10,X17,X14))) & ~(v1_xboole_0(k8_gate_1(X10,X17,X14)) & ~v1_xboole_0(X30)) & ~(v1_xboole_0(X29) & ~v1_xboole_0(k2_gate_1(X13,X17))) & ~(v1_xboole_0(k2_gate_1(X13,X17)) & ~v1_xboole_0(X29)) & ~(v1_xboole_0(X39) & ~v1_xboole_0(k4_gate_1(X28,X18))) & ~(v1_xboole_0(k4_gate_1(X28,X18)) & ~v1_xboole_0(X39)) & ~(v1_xboole_0(X28) & ~v1_xboole_0(k13_gate_1(X26,X27,X13))) & ~(v1_xboole_0(k13_gate_1(X26,X27,X13)) & ~v1_xboole_0(X28)) & ~(v1_xboole_0(X27) & ~v1_xboole_0(k8_gate_1(X14,X11,X22))) & ~(v1_xboole_0(k8_gate_1(X14,X11,X22)) & ~v1_xboole_0(X27)) & ~(v1_xboole_0(X26) & ~v1_xboole_0(k2_gate_1(X10,X14))) & ~(v1_xboole_0(k2_gate_1(X10,X14)) & ~v1_xboole_0(X26)) & ~(v1_xboole_0(X38) & ~v1_xboole_0(k4_gate_1(X25,X15))) & ~(v1_xboole_0(k4_gate_1(X25,X15)) & ~v1_xboole_0(X38)) & ~(v1_xboole_0(X25) & ~v1_xboole_0(k7_gate_1(X24,X10))) & ~(v1_xboole_0(k7_gate_1(X24,X10)) & ~v1_xboole_0(X25)) & ~(v1_xboole_0(X24) & ~v1_xboole_0(k2_gate_1(X22,X11))) & ~(v1_xboole_0(k2_gate_1(X22,X11)) & ~v1_xboole_0(X24)) & ~(v1_xboole_0(X37) & ~v1_xboole_0(k4_gate_1(X23,X12))) & ~(v1_xboole_0(k4_gate_1(X23,X12)) & ~v1_xboole_0(X37)) & ~(v1_xboole_0(X23) & ~v1_xboole_0(k1_gate_1(X22))) & ~(v1_xboole_0(k1_gate_1(X22)) & ~v1_xboole_0(X23)) & ~(v1_xboole_0(X22) & ~v1_xboole_0(k1_gate_1(X0))) & ~(v1_xboole_0(k1_gate_1(X0)) & ~v1_xboole_0(X22)) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k34_gate_1(X7,X8))) & ~(v1_xboole_0(k34_gate_1(X7,X8)) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k6_gate_1(X7,X8))) & ~(v1_xboole_0(k6_gate_1(X7,X8)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k7_gate_1(X7,X8))) & ~(v1_xboole_0(k7_gate_1(X7,X8)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X18) & ~v1_xboole_0(k34_gate_1(X5,X6))) & ~(v1_xboole_0(k34_gate_1(X5,X6)) & ~v1_xboole_0(X18)) & ~(v1_xboole_0(X17) & ~v1_xboole_0(k6_gate_1(X5,X6))) & ~(v1_xboole_0(k6_gate_1(X5,X6)) & ~v1_xboole_0(X17)) & ~(v1_xboole_0(X16) & ~v1_xboole_0(k7_gate_1(X5,X6))) & ~(v1_xboole_0(k7_gate_1(X5,X6)) & ~v1_xboole_0(X16)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k34_gate_1(X3,X4))) & ~(v1_xboole_0(k34_gate_1(X3,X4)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k6_gate_1(X3,X4))) & ~(v1_xboole_0(k6_gate_1(X3,X4)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k7_gate_1(X3,X4))) & ~(v1_xboole_0(k7_gate_1(X3,X4)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k34_gate_1(X1,X2))) & ~(v1_xboole_0(k34_gate_1(X1,X2)) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k6_gate_1(X1,X2))) & ~(v1_xboole_0(k6_gate_1(X1,X2)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k7_gate_1(X1,X2))) & ~(v1_xboole_0(k7_gate_1(X1,X2)) & ~v1_xboole_0(X10))) [flattening 2]
74. ! [X0,X1] : ((~(~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0))) => k34_gate_1(X0,X1) = k1_xboole_0) & ~(k34_gate_1(X0,X1) != k1_gate_1(k1_xboole_0) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0)))) [flattening 13]
75. ! [X0,X1,X2] : ((~(v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k13_gate_1(X0,X1,X2) = k1_xboole_0) & ~(k13_gate_1(X0,X1,X2) != k1_gate_1(k1_xboole_0) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [flattening 8]
76. ! [X0,X1,X2,X3] : ((~(~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k1_xboole_0 = k14_gate_1(X0,X1,X2,X3)) & ~(k1_gate_1(k1_xboole_0) != k14_gate_1(X0,X1,X2,X3) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 9]
77. ! [X0,X1,X2,X3] : ((~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k1_xboole_0 = k17_gate_1(X0,X1,X2,X3)) & ~(k1_gate_1(k1_xboole_0) != k17_gate_1(X0,X1,X2,X3) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [flattening 10]
78. ! [X0,X1,X2,X3,X4] : k36_gate_1(X0,X1,X2,X3,X4) = k11_gate_1(X2,X3,k11_gate_1(X0,X1,X4)) [flattening 15]
79. ! [X0,X1,X2,X3,X4] : k35_gate_1(X0,X1,X2,X3,X4) = k10_gate_1(X2,X3,k11_gate_1(X0,X1,X4)) [flattening 14]
80. ! [X0,X1,X2,X3,X4] : ((~(~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k1_xboole_0 = k18_gate_1(X0,X1,X2,X3,X4)) & ~(k1_gate_1(k1_xboole_0) != k18_gate_1(X0,X1,X2,X3,X4) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 11]
81. ! [X0,X1,X2,X3,X4] : ((~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k1_xboole_0 = k21_gate_1(X0,X1,X2,X3,X4)) & ~(k1_gate_1(k1_xboole_0) != k21_gate_1(X0,X1,X2,X3,X4) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [flattening 12]
82. ! [X0,X1,X2,X3,X4,X5,X6] : k38_gate_1(X0,X1,X2,X3,X4,X5,X6) = k11_gate_1(X4,X5,k36_gate_1(X0,X1,X2,X3,X6)) [flattening 17]
83. ! [X0,X1,X2,X3,X4,X5,X6] : k37_gate_1(X0,X1,X2,X3,X4,X5,X6) = k10_gate_1(X4,X5,k36_gate_1(X0,X1,X2,X3,X6)) [flattening 16]
84. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : k40_gate_1(X0,X1,X2,X3,X4,X5,X6,X7,X8) = k11_gate_1(X6,X7,k38_gate_1(X0,X1,X2,X3,X4,X5,X8)) [flattening 19]
85. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : k39_gate_1(X0,X1,X2,X3,X4,X5,X6,X7,X8) = k10_gate_1(X6,X7,k38_gate_1(X0,X1,X2,X3,X4,X5,X8)) [flattening 18]
86. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40] : (((v1_xboole_0(X9) & ~v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) | (v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X9)) | (v1_xboole_0(X40) & ~v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) | (v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X40)) | (v1_xboole_0(X39) & ~v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0))) | (v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0)) & ~v1_xboole_0(X39)) | (v1_xboole_0(X38) & ~v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0))) | (v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0)) & ~v1_xboole_0(X38)) | (v1_xboole_0(X37) & ~v1_xboole_0(k10_gate_1(X1,X2,X0))) | (v1_xboole_0(k10_gate_1(X1,X2,X0)) & ~v1_xboole_0(X37))) & (~v1_xboole_0(X9) | v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36))) & (~v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36)) | v1_xboole_0(X9)) & (~v1_xboole_0(X36) | v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22))) & (~v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22)) | v1_xboole_0(X36)) & (~v1_xboole_0(X35) | v1_xboole_0(k14_gate_1(X10,X20,X17,X14))) & (~v1_xboole_0(k14_gate_1(X10,X20,X17,X14)) | v1_xboole_0(X35)) & (~v1_xboole_0(X34) | v1_xboole_0(k8_gate_1(X13,X20,X17))) & (~v1_xboole_0(k8_gate_1(X13,X20,X17)) | v1_xboole_0(X34)) & (~v1_xboole_0(X33) | v1_xboole_0(k2_gate_1(X16,X20))) & (~v1_xboole_0(k2_gate_1(X16,X20)) | v1_xboole_0(X33)) & (~v1_xboole_0(X40) | v1_xboole_0(k4_gate_1(X32,X21))) & (~v1_xboole_0(k4_gate_1(X32,X21)) | v1_xboole_0(X40)) & (~v1_xboole_0(X32) | v1_xboole_0(k17_gate_1(X29,X30,X31,X16))) & (~v1_xboole_0(k17_gate_1(X29,X30,X31,X16)) | v1_xboole_0(X32)) & (~v1_xboole_0(X31) | v1_xboole_0(k14_gate_1(X17,X14,X11,X22))) & (~v1_xboole_0(k14_gate_1(X17,X14,X11,X22)) | v1_xboole_0(X31)) & (~v1_xboole_0(X30) | v1_xboole_0(k8_gate_1(X10,X17,X14))) & (~v1_xboole_0(k8_gate_1(X10,X17,X14)) | v1_xboole_0(X30)) & (~v1_xboole_0(X29) | v1_xboole_0(k2_gate_1(X13,X17))) & (~v1_xboole_0(k2_gate_1(X13,X17)) | v1_xboole_0(X29)) & (~v1_xboole_0(X39) | v1_xboole_0(k4_gate_1(X28,X18))) & (~v1_xboole_0(k4_gate_1(X28,X18)) | v1_xboole_0(X39)) & (~v1_xboole_0(X28) | v1_xboole_0(k13_gate_1(X26,X27,X13))) & (~v1_xboole_0(k13_gate_1(X26,X27,X13)) | v1_xboole_0(X28)) & (~v1_xboole_0(X27) | v1_xboole_0(k8_gate_1(X14,X11,X22))) & (~v1_xboole_0(k8_gate_1(X14,X11,X22)) | v1_xboole_0(X27)) & (~v1_xboole_0(X26) | v1_xboole_0(k2_gate_1(X10,X14))) & (~v1_xboole_0(k2_gate_1(X10,X14)) | v1_xboole_0(X26)) & (~v1_xboole_0(X38) | v1_xboole_0(k4_gate_1(X25,X15))) & (~v1_xboole_0(k4_gate_1(X25,X15)) | v1_xboole_0(X38)) & (~v1_xboole_0(X25) | v1_xboole_0(k7_gate_1(X24,X10))) & (~v1_xboole_0(k7_gate_1(X24,X10)) | v1_xboole_0(X25)) & (~v1_xboole_0(X24) | v1_xboole_0(k2_gate_1(X22,X11))) & (~v1_xboole_0(k2_gate_1(X22,X11)) | v1_xboole_0(X24)) & (~v1_xboole_0(X37) | v1_xboole_0(k4_gate_1(X23,X12))) & (~v1_xboole_0(k4_gate_1(X23,X12)) | v1_xboole_0(X37)) & (~v1_xboole_0(X23) | v1_xboole_0(k1_gate_1(X22))) & (~v1_xboole_0(k1_gate_1(X22)) | v1_xboole_0(X23)) & (~v1_xboole_0(X22) | v1_xboole_0(k1_gate_1(X0))) & (~v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X22)) & (~v1_xboole_0(X21) | v1_xboole_0(k34_gate_1(X7,X8))) & (~v1_xboole_0(k34_gate_1(X7,X8)) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(k6_gate_1(X7,X8))) & (~v1_xboole_0(k6_gate_1(X7,X8)) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k7_gate_1(X7,X8))) & (~v1_xboole_0(k7_gate_1(X7,X8)) | v1_xboole_0(X19)) & (~v1_xboole_0(X18) | v1_xboole_0(k34_gate_1(X5,X6))) & (~v1_xboole_0(k34_gate_1(X5,X6)) | v1_xboole_0(X18)) & (~v1_xboole_0(X17) | v1_xboole_0(k6_gate_1(X5,X6))) & (~v1_xboole_0(k6_gate_1(X5,X6)) | v1_xboole_0(X17)) & (~v1_xboole_0(X16) | v1_xboole_0(k7_gate_1(X5,X6))) & (~v1_xboole_0(k7_gate_1(X5,X6)) | v1_xboole_0(X16)) & (~v1_xboole_0(X15) | v1_xboole_0(k34_gate_1(X3,X4))) & (~v1_xboole_0(k34_gate_1(X3,X4)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k6_gate_1(X3,X4))) & (~v1_xboole_0(k6_gate_1(X3,X4)) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k7_gate_1(X3,X4))) & (~v1_xboole_0(k7_gate_1(X3,X4)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k34_gate_1(X1,X2))) & (~v1_xboole_0(k34_gate_1(X1,X2)) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k6_gate_1(X1,X2))) & (~v1_xboole_0(k6_gate_1(X1,X2)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k7_gate_1(X1,X2))) & (~v1_xboole_0(k7_gate_1(X1,X2)) | v1_xboole_0(X10))) [ennf transformation 73]
87. ! [X0] : (v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0)) [ennf transformation 56]
88. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 48]
89. ! [X0,X1] : ((k34_gate_1(X0,X1) = k1_xboole_0 | ((v1_xboole_0(X1) | v1_xboole_0(X0)) & (~v1_xboole_0(X1) | ~v1_xboole_0(X0)))) & (k34_gate_1(X0,X1) = k1_gate_1(k1_xboole_0) | (~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | (v1_xboole_0(X1) & v1_xboole_0(X0)))) [ennf transformation 74]
90. ! [X0,X1,X2] : ((k13_gate_1(X0,X1,X2) = k1_xboole_0 | (v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & (k13_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 75]
91. ! [X0,X1,X2,X3] : ((k1_xboole_0 = k14_gate_1(X0,X1,X2,X3) | (~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & (k1_gate_1(k1_xboole_0) = k14_gate_1(X0,X1,X2,X3) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 76]
92. ! [X0,X1,X2,X3] : ((k1_xboole_0 = k17_gate_1(X0,X1,X2,X3) | (v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & (k1_gate_1(k1_xboole_0) = k17_gate_1(X0,X1,X2,X3) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 77]
93. ! [X0,X1,X2,X3,X4] : ((k1_xboole_0 = k18_gate_1(X0,X1,X2,X3,X4) | (~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & (k1_gate_1(k1_xboole_0) = k18_gate_1(X0,X1,X2,X3,X4) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 80]
94. ! [X0,X1,X2,X3,X4] : ((k1_xboole_0 = k21_gate_1(X0,X1,X2,X3,X4) | (v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & (k1_gate_1(k1_xboole_0) = k21_gate_1(X0,X1,X2,X3,X4) | ~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 81]
95. ! [X0,X1] : (v1_xboole_0(k7_gate_1(X0,X1)) | v1_xboole_0(X0)) [ennf transformation 44]
96. ! [X0,X1] : (~v1_xboole_0(k6_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 41]
97. ! [X0,X1] : (v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 67]
98. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 66]
99. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 98]
100. ! [X0,X1] : (v1_xboole_0(k6_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 42]
101. ! [X0,X1] : (v1_xboole_0(k6_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 100]
102. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 70]
103. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 102]
104. ! [X0,X1] : (~v1_xboole_0(k4_gate_1(X0,X1)) | (v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 68]
105. ! [X0,X1] : (~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 104]
106. ! [X0,X1] : (~v1_xboole_0(k7_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 43]
107. ! [X0,X1] : (~v1_xboole_0(k7_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 106]
108. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 69]
109. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 108]
110. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0)) [ennf transformation 47]
111. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0)) [ennf transformation 50]
112. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0)) [ennf transformation 46]
113. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 45]
114. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 113]
115. ! [X0,X1,X2] : (~v1_xboole_0(k10_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 59]
116. ! [X0,X1,X2] : (~v1_xboole_0(k10_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 115]
117. ! [X0,X1,X2] : (v1_xboole_0(k10_gate_1(X2,X0,X1)) | (~v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 58]
118. ! [X0,X1,X2] : (v1_xboole_0(k10_gate_1(X2,X0,X1)) | ~v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 117]
119. ! [X0,X1,X2] : (v1_xboole_0(k10_gate_1(X0,X1,X2)) | (~v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 55]
120. ! [X0,X1,X2] : (v1_xboole_0(k10_gate_1(X0,X1,X2)) | ~v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 119]
121. ! [X0,X1,X2] : (v1_xboole_0(k10_gate_1(X0,X2,X1)) | (~v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 57]
122. ! [X0,X1,X2] : (v1_xboole_0(k10_gate_1(X0,X2,X1)) | ~v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 121]
123. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X2,X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 62]
124. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X2,X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 123]
125. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X0,X1,X2)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 60]
126. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X0,X1,X2)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 125]
127. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X0,X2,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 61]
128. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X0,X2,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 127]
129. ! [X0,X1,X2] : (~v1_xboole_0(k10_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 52]
130. ! [X0,X1,X2] : (~v1_xboole_0(k10_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 129]
131. ! [X0,X1,X2] : (~v1_xboole_0(k10_gate_1(X0,X2,X1)) | (v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 53]
132. ! [X0,X1,X2] : (~v1_xboole_0(k10_gate_1(X0,X2,X1)) | v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 131]
133. ! [X0,X1,X2] : (~v1_xboole_0(k10_gate_1(X2,X0,X1)) | (v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 54]
134. ! [X0,X1,X2] : (~v1_xboole_0(k10_gate_1(X2,X0,X1)) | v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 133]
135. ! [X0,X1,X2] : (v1_xboole_0(k10_gate_1(X0,X1,X2)) | (~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 51]
136. ! [X0,X1,X2] : (v1_xboole_0(k10_gate_1(X0,X1,X2)) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 135]
137. ! [X0,X1,X2] : (v1_xboole_0(k11_gate_1(X0,X1,X2)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 63]
138. ! [X0,X1,X2] : (v1_xboole_0(k11_gate_1(X0,X1,X2)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 137]
139. ! [X0,X1,X2] : (v1_xboole_0(k11_gate_1(X0,X2,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 64]
140. ! [X0,X1,X2] : (v1_xboole_0(k11_gate_1(X0,X2,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 139]
141. ! [X0,X1,X2] : (v1_xboole_0(k11_gate_1(X2,X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 65]
142. ! [X0,X1,X2] : (v1_xboole_0(k11_gate_1(X2,X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 141]
143. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40] : (((v1_xboole_0(X9) & ~v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) | (v1_xboole_0(k40_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X9)) | (v1_xboole_0(X40) & ~v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0))) | (v1_xboole_0(k39_gate_1(X1,X2,X3,X4,X5,X6,X7,X8,X0)) & ~v1_xboole_0(X40)) | (v1_xboole_0(X39) & ~v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0))) | (v1_xboole_0(k37_gate_1(X1,X2,X3,X4,X5,X6,X0)) & ~v1_xboole_0(X39)) | (v1_xboole_0(X38) & ~v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0))) | (v1_xboole_0(k35_gate_1(X1,X2,X3,X4,X0)) & ~v1_xboole_0(X38)) | (v1_xboole_0(X37) & ~v1_xboole_0(k10_gate_1(X1,X2,X0))) | (v1_xboole_0(k10_gate_1(X1,X2,X0)) & ~v1_xboole_0(X37))) & (~v1_xboole_0(X9) | v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36))) & (~v1_xboole_0(k21_gate_1(X19,X33,X34,X35,X36)) | v1_xboole_0(X9)) & (~v1_xboole_0(X36) | v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22))) & (~v1_xboole_0(k18_gate_1(X20,X17,X14,X11,X22)) | v1_xboole_0(X36)) & (~v1_xboole_0(X35) | v1_xboole_0(k14_gate_1(X10,X20,X17,X14))) & (~v1_xboole_0(k14_gate_1(X10,X20,X17,X14)) | v1_xboole_0(X35)) & (~v1_xboole_0(X34) | v1_xboole_0(k8_gate_1(X13,X20,X17))) & (~v1_xboole_0(k8_gate_1(X13,X20,X17)) | v1_xboole_0(X34)) & (~v1_xboole_0(X33) | v1_xboole_0(k2_gate_1(X16,X20))) & (~v1_xboole_0(k2_gate_1(X16,X20)) | v1_xboole_0(X33)) & (~v1_xboole_0(X40) | v1_xboole_0(k4_gate_1(X32,X21))) & (~v1_xboole_0(k4_gate_1(X32,X21)) | v1_xboole_0(X40)) & (~v1_xboole_0(X32) | v1_xboole_0(k17_gate_1(X29,X30,X31,X16))) & (~v1_xboole_0(k17_gate_1(X29,X30,X31,X16)) | v1_xboole_0(X32)) & (~v1_xboole_0(X31) | v1_xboole_0(k14_gate_1(X17,X14,X11,X22))) & (~v1_xboole_0(k14_gate_1(X17,X14,X11,X22)) | v1_xboole_0(X31)) & (~v1_xboole_0(X30) | v1_xboole_0(k8_gate_1(X10,X17,X14))) & (~v1_xboole_0(k8_gate_1(X10,X17,X14)) | v1_xboole_0(X30)) & (~v1_xboole_0(X29) | v1_xboole_0(k2_gate_1(X13,X17))) & (~v1_xboole_0(k2_gate_1(X13,X17)) | v1_xboole_0(X29)) & (~v1_xboole_0(X39) | v1_xboole_0(k4_gate_1(X28,X18))) & (~v1_xboole_0(k4_gate_1(X28,X18)) | v1_xboole_0(X39)) & (~v1_xboole_0(X28) | v1_xboole_0(k13_gate_1(X26,X27,X13))) & (~v1_xboole_0(k13_gate_1(X26,X27,X13)) | v1_xboole_0(X28)) & (~v1_xboole_0(X27) | v1_xboole_0(k8_gate_1(X14,X11,X22))) & (~v1_xboole_0(k8_gate_1(X14,X11,X22)) | v1_xboole_0(X27)) & (~v1_xboole_0(X26) | v1_xboole_0(k2_gate_1(X10,X14))) & (~v1_xboole_0(k2_gate_1(X10,X14)) | v1_xboole_0(X26)) & (~v1_xboole_0(X38) | v1_xboole_0(k4_gate_1(X25,X15))) & (~v1_xboole_0(k4_gate_1(X25,X15)) | v1_xboole_0(X38)) & (~v1_xboole_0(X25) | v1_xboole_0(k7_gate_1(X24,X10))) & (~v1_xboole_0(k7_gate_1(X24,X10)) | v1_xboole_0(X25)) & (~v1_xboole_0(X24) | v1_xboole_0(k2_gate_1(X22,X11))) & (~v1_xboole_0(k2_gate_1(X22,X11)) | v1_xboole_0(X24)) & (~v1_xboole_0(X37) | v1_xboole_0(k4_gate_1(X23,X12))) & (~v1_xboole_0(k4_gate_1(X23,X12)) | v1_xboole_0(X37)) & (~v1_xboole_0(X23) | v1_xboole_0(k1_gate_1(X22))) & (~v1_xboole_0(k1_gate_1(X22)) | v1_xboole_0(X23)) & (~v1_xboole_0(X22) | v1_xboole_0(k1_gate_1(X0))) & (~v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X22)) & (~v1_xboole_0(X21) | v1_xboole_0(k34_gate_1(X7,X8))) & (~v1_xboole_0(k34_gate_1(X7,X8)) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(k6_gate_1(X7,X8))) & (~v1_xboole_0(k6_gate_1(X7,X8)) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k7_gate_1(X7,X8))) & (~v1_xboole_0(k7_gate_1(X7,X8)) | v1_xboole_0(X19)) & (~v1_xboole_0(X18) | v1_xboole_0(k34_gate_1(X5,X6))) & (~v1_xboole_0(k34_gate_1(X5,X6)) | v1_xboole_0(X18)) & (~v1_xboole_0(X17) | v1_xboole_0(k6_gate_1(X5,X6))) & (~v1_xboole_0(k6_gate_1(X5,X6)) | v1_xboole_0(X17)) & (~v1_xboole_0(X16) | v1_xboole_0(k7_gate_1(X5,X6))) & (~v1_xboole_0(k7_gate_1(X5,X6)) | v1_xboole_0(X16)) & (~v1_xboole_0(X15) | v1_xboole_0(k34_gate_1(X3,X4))) & (~v1_xboole_0(k34_gate_1(X3,X4)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k6_gate_1(X3,X4))) & (~v1_xboole_0(k6_gate_1(X3,X4)) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k7_gate_1(X3,X4))) & (~v1_xboole_0(k7_gate_1(X3,X4)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k34_gate_1(X1,X2))) & (~v1_xboole_0(k34_gate_1(X1,X2)) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k6_gate_1(X1,X2))) & (~v1_xboole_0(k6_gate_1(X1,X2)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k7_gate_1(X1,X2))) & (~v1_xboole_0(k7_gate_1(X1,X2)) | v1_xboole_0(X10))) => (((v1_xboole_0(sK9) & ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0))) | (v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) & ~v1_xboole_0(sK9)) | (v1_xboole_0(sK40) & ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0))) | (v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) & ~v1_xboole_0(sK40)) | (v1_xboole_0(sK39) & ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0))) | (v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) & ~v1_xboole_0(sK39)) | (v1_xboole_0(sK38) & ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0))) | (v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) & ~v1_xboole_0(sK38)) | (v1_xboole_0(sK37) & ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0))) | (v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) & ~v1_xboole_0(sK37))) & (~v1_xboole_0(sK9) | v1_xboole_0(k21_gate_1(sK19,sK33,sK34,sK35,sK36))) & (~v1_xboole_0(k21_gate_1(sK19,sK33,sK34,sK35,sK36)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK36) | v1_xboole_0(k18_gate_1(sK20,sK17,sK14,sK11,sK22))) & (~v1_xboole_0(k18_gate_1(sK20,sK17,sK14,sK11,sK22)) | v1_xboole_0(sK36)) & (~v1_xboole_0(sK35) | v1_xboole_0(k14_gate_1(sK10,sK20,sK17,sK14))) & (~v1_xboole_0(k14_gate_1(sK10,sK20,sK17,sK14)) | v1_xboole_0(sK35)) & (~v1_xboole_0(sK34) | v1_xboole_0(k8_gate_1(sK13,sK20,sK17))) & (~v1_xboole_0(k8_gate_1(sK13,sK20,sK17)) | v1_xboole_0(sK34)) & (~v1_xboole_0(sK33) | v1_xboole_0(k2_gate_1(sK16,sK20))) & (~v1_xboole_0(k2_gate_1(sK16,sK20)) | v1_xboole_0(sK33)) & (~v1_xboole_0(sK40) | v1_xboole_0(k4_gate_1(sK32,sK21))) & (~v1_xboole_0(k4_gate_1(sK32,sK21)) | v1_xboole_0(sK40)) & (~v1_xboole_0(sK32) | v1_xboole_0(k17_gate_1(sK29,sK30,sK31,sK16))) & (~v1_xboole_0(k17_gate_1(sK29,sK30,sK31,sK16)) | v1_xboole_0(sK32)) & (~v1_xboole_0(sK31) | v1_xboole_0(k14_gate_1(sK17,sK14,sK11,sK22))) & (~v1_xboole_0(k14_gate_1(sK17,sK14,sK11,sK22)) | v1_xboole_0(sK31)) & (~v1_xboole_0(sK30) | v1_xboole_0(k8_gate_1(sK10,sK17,sK14))) & (~v1_xboole_0(k8_gate_1(sK10,sK17,sK14)) | v1_xboole_0(sK30)) & (~v1_xboole_0(sK29) | v1_xboole_0(k2_gate_1(sK13,sK17))) & (~v1_xboole_0(k2_gate_1(sK13,sK17)) | v1_xboole_0(sK29)) & (~v1_xboole_0(sK39) | v1_xboole_0(k4_gate_1(sK28,sK18))) & (~v1_xboole_0(k4_gate_1(sK28,sK18)) | v1_xboole_0(sK39)) & (~v1_xboole_0(sK28) | v1_xboole_0(k13_gate_1(sK26,sK27,sK13))) & (~v1_xboole_0(k13_gate_1(sK26,sK27,sK13)) | v1_xboole_0(sK28)) & (~v1_xboole_0(sK27) | v1_xboole_0(k8_gate_1(sK14,sK11,sK22))) & (~v1_xboole_0(k8_gate_1(sK14,sK11,sK22)) | v1_xboole_0(sK27)) & (~v1_xboole_0(sK26) | v1_xboole_0(k2_gate_1(sK10,sK14))) & (~v1_xboole_0(k2_gate_1(sK10,sK14)) | v1_xboole_0(sK26)) & (~v1_xboole_0(sK38) | v1_xboole_0(k4_gate_1(sK25,sK15))) & (~v1_xboole_0(k4_gate_1(sK25,sK15)) | v1_xboole_0(sK38)) & (~v1_xboole_0(sK25) | v1_xboole_0(k7_gate_1(sK24,sK10))) & (~v1_xboole_0(k7_gate_1(sK24,sK10)) | v1_xboole_0(sK25)) & (~v1_xboole_0(sK24) | v1_xboole_0(k2_gate_1(sK22,sK11))) & (~v1_xboole_0(k2_gate_1(sK22,sK11)) | v1_xboole_0(sK24)) & (~v1_xboole_0(sK37) | v1_xboole_0(k4_gate_1(sK23,sK12))) & (~v1_xboole_0(k4_gate_1(sK23,sK12)) | v1_xboole_0(sK37)) & (~v1_xboole_0(sK23) | v1_xboole_0(k1_gate_1(sK22))) & (~v1_xboole_0(k1_gate_1(sK22)) | v1_xboole_0(sK23)) & (~v1_xboole_0(sK22) | v1_xboole_0(k1_gate_1(sK0))) & (~v1_xboole_0(k1_gate_1(sK0)) | v1_xboole_0(sK22)) & (~v1_xboole_0(sK21) | v1_xboole_0(k34_gate_1(sK7,sK8))) & (~v1_xboole_0(k34_gate_1(sK7,sK8)) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(k6_gate_1(sK7,sK8))) & (~v1_xboole_0(k6_gate_1(sK7,sK8)) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k7_gate_1(sK7,sK8))) & (~v1_xboole_0(k7_gate_1(sK7,sK8)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK18) | v1_xboole_0(k34_gate_1(sK5,sK6))) & (~v1_xboole_0(k34_gate_1(sK5,sK6)) | v1_xboole_0(sK18)) & (~v1_xboole_0(sK17) | v1_xboole_0(k6_gate_1(sK5,sK6))) & (~v1_xboole_0(k6_gate_1(sK5,sK6)) | v1_xboole_0(sK17)) & (~v1_xboole_0(sK16) | v1_xboole_0(k7_gate_1(sK5,sK6))) & (~v1_xboole_0(k7_gate_1(sK5,sK6)) | v1_xboole_0(sK16)) & (~v1_xboole_0(sK15) | v1_xboole_0(k34_gate_1(sK3,sK4))) & (~v1_xboole_0(k34_gate_1(sK3,sK4)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k6_gate_1(sK3,sK4))) & (~v1_xboole_0(k6_gate_1(sK3,sK4)) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k7_gate_1(sK3,sK4))) & (~v1_xboole_0(k7_gate_1(sK3,sK4)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k34_gate_1(sK1,sK2))) & (~v1_xboole_0(k34_gate_1(sK1,sK2)) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k6_gate_1(sK1,sK2))) & (~v1_xboole_0(k6_gate_1(sK1,sK2)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k7_gate_1(sK1,sK2))) & (~v1_xboole_0(k7_gate_1(sK1,sK2)) | v1_xboole_0(sK10))) [choice axiom]
144. ((v1_xboole_0(sK9) & ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0))) | (v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) & ~v1_xboole_0(sK9)) | (v1_xboole_0(sK40) & ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0))) | (v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) & ~v1_xboole_0(sK40)) | (v1_xboole_0(sK39) & ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0))) | (v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) & ~v1_xboole_0(sK39)) | (v1_xboole_0(sK38) & ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0))) | (v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) & ~v1_xboole_0(sK38)) | (v1_xboole_0(sK37) & ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0))) | (v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) & ~v1_xboole_0(sK37))) & (~v1_xboole_0(sK9) | v1_xboole_0(k21_gate_1(sK19,sK33,sK34,sK35,sK36))) & (~v1_xboole_0(k21_gate_1(sK19,sK33,sK34,sK35,sK36)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK36) | v1_xboole_0(k18_gate_1(sK20,sK17,sK14,sK11,sK22))) & (~v1_xboole_0(k18_gate_1(sK20,sK17,sK14,sK11,sK22)) | v1_xboole_0(sK36)) & (~v1_xboole_0(sK35) | v1_xboole_0(k14_gate_1(sK10,sK20,sK17,sK14))) & (~v1_xboole_0(k14_gate_1(sK10,sK20,sK17,sK14)) | v1_xboole_0(sK35)) & (~v1_xboole_0(sK34) | v1_xboole_0(k8_gate_1(sK13,sK20,sK17))) & (~v1_xboole_0(k8_gate_1(sK13,sK20,sK17)) | v1_xboole_0(sK34)) & (~v1_xboole_0(sK33) | v1_xboole_0(k2_gate_1(sK16,sK20))) & (~v1_xboole_0(k2_gate_1(sK16,sK20)) | v1_xboole_0(sK33)) & (~v1_xboole_0(sK40) | v1_xboole_0(k4_gate_1(sK32,sK21))) & (~v1_xboole_0(k4_gate_1(sK32,sK21)) | v1_xboole_0(sK40)) & (~v1_xboole_0(sK32) | v1_xboole_0(k17_gate_1(sK29,sK30,sK31,sK16))) & (~v1_xboole_0(k17_gate_1(sK29,sK30,sK31,sK16)) | v1_xboole_0(sK32)) & (~v1_xboole_0(sK31) | v1_xboole_0(k14_gate_1(sK17,sK14,sK11,sK22))) & (~v1_xboole_0(k14_gate_1(sK17,sK14,sK11,sK22)) | v1_xboole_0(sK31)) & (~v1_xboole_0(sK30) | v1_xboole_0(k8_gate_1(sK10,sK17,sK14))) & (~v1_xboole_0(k8_gate_1(sK10,sK17,sK14)) | v1_xboole_0(sK30)) & (~v1_xboole_0(sK29) | v1_xboole_0(k2_gate_1(sK13,sK17))) & (~v1_xboole_0(k2_gate_1(sK13,sK17)) | v1_xboole_0(sK29)) & (~v1_xboole_0(sK39) | v1_xboole_0(k4_gate_1(sK28,sK18))) & (~v1_xboole_0(k4_gate_1(sK28,sK18)) | v1_xboole_0(sK39)) & (~v1_xboole_0(sK28) | v1_xboole_0(k13_gate_1(sK26,sK27,sK13))) & (~v1_xboole_0(k13_gate_1(sK26,sK27,sK13)) | v1_xboole_0(sK28)) & (~v1_xboole_0(sK27) | v1_xboole_0(k8_gate_1(sK14,sK11,sK22))) & (~v1_xboole_0(k8_gate_1(sK14,sK11,sK22)) | v1_xboole_0(sK27)) & (~v1_xboole_0(sK26) | v1_xboole_0(k2_gate_1(sK10,sK14))) & (~v1_xboole_0(k2_gate_1(sK10,sK14)) | v1_xboole_0(sK26)) & (~v1_xboole_0(sK38) | v1_xboole_0(k4_gate_1(sK25,sK15))) & (~v1_xboole_0(k4_gate_1(sK25,sK15)) | v1_xboole_0(sK38)) & (~v1_xboole_0(sK25) | v1_xboole_0(k7_gate_1(sK24,sK10))) & (~v1_xboole_0(k7_gate_1(sK24,sK10)) | v1_xboole_0(sK25)) & (~v1_xboole_0(sK24) | v1_xboole_0(k2_gate_1(sK22,sK11))) & (~v1_xboole_0(k2_gate_1(sK22,sK11)) | v1_xboole_0(sK24)) & (~v1_xboole_0(sK37) | v1_xboole_0(k4_gate_1(sK23,sK12))) & (~v1_xboole_0(k4_gate_1(sK23,sK12)) | v1_xboole_0(sK37)) & (~v1_xboole_0(sK23) | v1_xboole_0(k1_gate_1(sK22))) & (~v1_xboole_0(k1_gate_1(sK22)) | v1_xboole_0(sK23)) & (~v1_xboole_0(sK22) | v1_xboole_0(k1_gate_1(sK0))) & (~v1_xboole_0(k1_gate_1(sK0)) | v1_xboole_0(sK22)) & (~v1_xboole_0(sK21) | v1_xboole_0(k34_gate_1(sK7,sK8))) & (~v1_xboole_0(k34_gate_1(sK7,sK8)) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(k6_gate_1(sK7,sK8))) & (~v1_xboole_0(k6_gate_1(sK7,sK8)) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k7_gate_1(sK7,sK8))) & (~v1_xboole_0(k7_gate_1(sK7,sK8)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK18) | v1_xboole_0(k34_gate_1(sK5,sK6))) & (~v1_xboole_0(k34_gate_1(sK5,sK6)) | v1_xboole_0(sK18)) & (~v1_xboole_0(sK17) | v1_xboole_0(k6_gate_1(sK5,sK6))) & (~v1_xboole_0(k6_gate_1(sK5,sK6)) | v1_xboole_0(sK17)) & (~v1_xboole_0(sK16) | v1_xboole_0(k7_gate_1(sK5,sK6))) & (~v1_xboole_0(k7_gate_1(sK5,sK6)) | v1_xboole_0(sK16)) & (~v1_xboole_0(sK15) | v1_xboole_0(k34_gate_1(sK3,sK4))) & (~v1_xboole_0(k34_gate_1(sK3,sK4)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k6_gate_1(sK3,sK4))) & (~v1_xboole_0(k6_gate_1(sK3,sK4)) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k7_gate_1(sK3,sK4))) & (~v1_xboole_0(k7_gate_1(sK3,sK4)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k34_gate_1(sK1,sK2))) & (~v1_xboole_0(k34_gate_1(sK1,sK2)) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k6_gate_1(sK1,sK2))) & (~v1_xboole_0(k6_gate_1(sK1,sK2)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k7_gate_1(sK1,sK2))) & (~v1_xboole_0(k7_gate_1(sK1,sK2)) | v1_xboole_0(sK10)) [skolemisation 86,143]
149. ~v1_xboole_0(k7_gate_1(sK1,sK2)) | v1_xboole_0(sK10) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
150. ~v1_xboole_0(sK10) | v1_xboole_0(k7_gate_1(sK1,sK2)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
151. ~v1_xboole_0(k6_gate_1(sK1,sK2)) | v1_xboole_0(sK11) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
152. ~v1_xboole_0(sK11) | v1_xboole_0(k6_gate_1(sK1,sK2)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
153. ~v1_xboole_0(k34_gate_1(sK1,sK2)) | v1_xboole_0(sK12) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
154. ~v1_xboole_0(sK12) | v1_xboole_0(k34_gate_1(sK1,sK2)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
155. ~v1_xboole_0(k7_gate_1(sK3,sK4)) | v1_xboole_0(sK13) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
156. ~v1_xboole_0(sK13) | v1_xboole_0(k7_gate_1(sK3,sK4)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
157. ~v1_xboole_0(k6_gate_1(sK3,sK4)) | v1_xboole_0(sK14) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
158. ~v1_xboole_0(sK14) | v1_xboole_0(k6_gate_1(sK3,sK4)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
159. ~v1_xboole_0(k34_gate_1(sK3,sK4)) | v1_xboole_0(sK15) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
160. ~v1_xboole_0(sK15) | v1_xboole_0(k34_gate_1(sK3,sK4)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
161. ~v1_xboole_0(k7_gate_1(sK5,sK6)) | v1_xboole_0(sK16) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
162. ~v1_xboole_0(sK16) | v1_xboole_0(k7_gate_1(sK5,sK6)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
163. ~v1_xboole_0(k6_gate_1(sK5,sK6)) | v1_xboole_0(sK17) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
164. ~v1_xboole_0(sK17) | v1_xboole_0(k6_gate_1(sK5,sK6)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
165. ~v1_xboole_0(k34_gate_1(sK5,sK6)) | v1_xboole_0(sK18) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
166. ~v1_xboole_0(sK18) | v1_xboole_0(k34_gate_1(sK5,sK6)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
167. ~v1_xboole_0(k7_gate_1(sK7,sK8)) | v1_xboole_0(sK19) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
168. ~v1_xboole_0(sK19) | v1_xboole_0(k7_gate_1(sK7,sK8)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
169. ~v1_xboole_0(k6_gate_1(sK7,sK8)) | v1_xboole_0(sK20) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
170. ~v1_xboole_0(sK20) | v1_xboole_0(k6_gate_1(sK7,sK8)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
171. ~v1_xboole_0(k34_gate_1(sK7,sK8)) | v1_xboole_0(sK21) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
172. ~v1_xboole_0(sK21) | v1_xboole_0(k34_gate_1(sK7,sK8)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
173. ~v1_xboole_0(k1_gate_1(sK0)) | v1_xboole_0(sK22) [cnf transformation 144] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
174. ~v1_xboole_0(sK22) | v1_xboole_0(k1_gate_1(sK0)) [cnf transformation 144] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
175. ~v1_xboole_0(k1_gate_1(sK22)) | v1_xboole_0(sK23) [cnf transformation 144] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
176. ~v1_xboole_0(sK23) | v1_xboole_0(k1_gate_1(sK22)) [cnf transformation 144] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
177. ~v1_xboole_0(k4_gate_1(sK23,sK12)) | v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
178. ~v1_xboole_0(sK37) | v1_xboole_0(k4_gate_1(sK23,sK12)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
179. ~v1_xboole_0(k2_gate_1(sK22,sK11)) | v1_xboole_0(sK24) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
180. ~v1_xboole_0(sK24) | v1_xboole_0(k2_gate_1(sK22,sK11)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
181. ~v1_xboole_0(k7_gate_1(sK24,sK10)) | v1_xboole_0(sK25) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
182. ~v1_xboole_0(sK25) | v1_xboole_0(k7_gate_1(sK24,sK10)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
183. ~v1_xboole_0(k4_gate_1(sK25,sK15)) | v1_xboole_0(sK38) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
184. ~v1_xboole_0(sK38) | v1_xboole_0(k4_gate_1(sK25,sK15)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
185. ~v1_xboole_0(k2_gate_1(sK10,sK14)) | v1_xboole_0(sK26) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
186. ~v1_xboole_0(sK26) | v1_xboole_0(k2_gate_1(sK10,sK14)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
187. ~v1_xboole_0(k8_gate_1(sK14,sK11,sK22)) | v1_xboole_0(sK27) [cnf transformation 144] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
188. ~v1_xboole_0(sK27) | v1_xboole_0(k8_gate_1(sK14,sK11,sK22)) [cnf transformation 144] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
189. ~v1_xboole_0(k13_gate_1(sK26,sK27,sK13)) | v1_xboole_0(sK28) [cnf transformation 144] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
190. ~v1_xboole_0(sK28) | v1_xboole_0(k13_gate_1(sK26,sK27,sK13)) [cnf transformation 144] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
191. ~v1_xboole_0(k4_gate_1(sK28,sK18)) | v1_xboole_0(sK39) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
192. ~v1_xboole_0(sK39) | v1_xboole_0(k4_gate_1(sK28,sK18)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
193. ~v1_xboole_0(k2_gate_1(sK13,sK17)) | v1_xboole_0(sK29) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
194. ~v1_xboole_0(sK29) | v1_xboole_0(k2_gate_1(sK13,sK17)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
195. ~v1_xboole_0(k8_gate_1(sK10,sK17,sK14)) | v1_xboole_0(sK30) [cnf transformation 144] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
196. ~v1_xboole_0(sK30) | v1_xboole_0(k8_gate_1(sK10,sK17,sK14)) [cnf transformation 144] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
197. ~v1_xboole_0(k14_gate_1(sK17,sK14,sK11,sK22)) | v1_xboole_0(sK31) [cnf transformation 144] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
198. ~v1_xboole_0(sK31) | v1_xboole_0(k14_gate_1(sK17,sK14,sK11,sK22)) [cnf transformation 144] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
199. ~v1_xboole_0(k17_gate_1(sK29,sK30,sK31,sK16)) | v1_xboole_0(sK32) [cnf transformation 144] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
200. ~v1_xboole_0(sK32) | v1_xboole_0(k17_gate_1(sK29,sK30,sK31,sK16)) [cnf transformation 144] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
201. ~v1_xboole_0(k4_gate_1(sK32,sK21)) | v1_xboole_0(sK40) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
202. ~v1_xboole_0(sK40) | v1_xboole_0(k4_gate_1(sK32,sK21)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
203. ~v1_xboole_0(k2_gate_1(sK16,sK20)) | v1_xboole_0(sK33) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
204. ~v1_xboole_0(sK33) | v1_xboole_0(k2_gate_1(sK16,sK20)) [cnf transformation 144] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
205. ~v1_xboole_0(k8_gate_1(sK13,sK20,sK17)) | v1_xboole_0(sK34) [cnf transformation 144] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
206. ~v1_xboole_0(sK34) | v1_xboole_0(k8_gate_1(sK13,sK20,sK17)) [cnf transformation 144] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
207. ~v1_xboole_0(k14_gate_1(sK10,sK20,sK17,sK14)) | v1_xboole_0(sK35) [cnf transformation 144] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
208. ~v1_xboole_0(sK35) | v1_xboole_0(k14_gate_1(sK10,sK20,sK17,sK14)) [cnf transformation 144] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
209. ~v1_xboole_0(k18_gate_1(sK20,sK17,sK14,sK11,sK22)) | v1_xboole_0(sK36) [cnf transformation 144] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
210. ~v1_xboole_0(sK36) | v1_xboole_0(k18_gate_1(sK20,sK17,sK14,sK11,sK22)) [cnf transformation 144] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
211. ~v1_xboole_0(k21_gate_1(sK19,sK33,sK34,sK35,sK36)) | v1_xboole_0(sK9) [cnf transformation 144] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
212. ~v1_xboole_0(sK9) | v1_xboole_0(k21_gate_1(sK19,sK33,sK34,sK35,sK36)) [cnf transformation 144] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
213. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
216. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
225. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
228. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
261. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
264. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
273. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
276. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
405. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
408. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
417. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
420. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
453. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
456. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
465. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
468. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK9) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
981. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
984. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
993. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
996. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1029. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1032. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1041. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1044. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(sK40) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1173. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1176. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1185. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1188. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK39) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1221. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1224. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK38) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1233. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK37) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1236. v1_xboole_0(sK9) | v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK40) | v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) | v1_xboole_0(sK39) | v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK38) | v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK37) | v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [cnf transformation 144] {a:0,w:53,wCS:5300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1237. v1_xboole_0(k1_xboole_0) [cnf transformation 49] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
1238. v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0) [cnf transformation 87] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1239. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 88] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1241. k34_gate_1(X0,X1) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X0) | v1_xboole_0(X1) [cnf transformation 89] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1242. k34_gate_1(X0,X1) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 89] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1244. k34_gate_1(X0,X1) = k1_xboole_0 | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 89] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1245. k34_gate_1(X0,X1) = k1_xboole_0 | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 89] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1246. k13_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 90] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1247. k13_gate_1(X0,X1,X2) = k1_xboole_0 | v1_xboole_0(X0) [cnf transformation 90] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1248. k13_gate_1(X0,X1,X2) = k1_xboole_0 | v1_xboole_0(X1) [cnf transformation 90] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1249. k13_gate_1(X0,X1,X2) = k1_xboole_0 | v1_xboole_0(X2) [cnf transformation 90] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1250. k1_gate_1(k1_xboole_0) = k14_gate_1(X0,X1,X2,X3) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 91] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1251. k1_xboole_0 = k14_gate_1(X0,X1,X2,X3) | ~v1_xboole_0(X0) [cnf transformation 91] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1252. k1_xboole_0 = k14_gate_1(X0,X1,X2,X3) | ~v1_xboole_0(X1) [cnf transformation 91] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1253. k1_xboole_0 = k14_gate_1(X0,X1,X2,X3) | ~v1_xboole_0(X2) [cnf transformation 91] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1254. k1_xboole_0 = k14_gate_1(X0,X1,X2,X3) | ~v1_xboole_0(X3) [cnf transformation 91] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1255. k1_gate_1(k1_xboole_0) = k17_gate_1(X0,X1,X2,X3) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 92] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1256. k1_xboole_0 = k17_gate_1(X0,X1,X2,X3) | v1_xboole_0(X0) [cnf transformation 92] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1257. k1_xboole_0 = k17_gate_1(X0,X1,X2,X3) | v1_xboole_0(X1) [cnf transformation 92] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1258. k1_xboole_0 = k17_gate_1(X0,X1,X2,X3) | v1_xboole_0(X2) [cnf transformation 92] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1259. k1_xboole_0 = k17_gate_1(X0,X1,X2,X3) | v1_xboole_0(X3) [cnf transformation 92] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1260. k36_gate_1(X0,X1,X2,X3,X4) = k11_gate_1(X2,X3,k11_gate_1(X0,X1,X4)) [cnf transformation 78] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1261. k35_gate_1(X0,X1,X2,X3,X4) = k10_gate_1(X2,X3,k11_gate_1(X0,X1,X4)) [cnf transformation 79] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1262. k1_gate_1(k1_xboole_0) = k18_gate_1(X0,X1,X2,X3,X4) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 93] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1263. k1_xboole_0 = k18_gate_1(X0,X1,X2,X3,X4) | ~v1_xboole_0(X0) [cnf transformation 93] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1264. k1_xboole_0 = k18_gate_1(X0,X1,X2,X3,X4) | ~v1_xboole_0(X1) [cnf transformation 93] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1265. k1_xboole_0 = k18_gate_1(X0,X1,X2,X3,X4) | ~v1_xboole_0(X2) [cnf transformation 93] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1266. k1_xboole_0 = k18_gate_1(X0,X1,X2,X3,X4) | ~v1_xboole_0(X3) [cnf transformation 93] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1267. k1_xboole_0 = k18_gate_1(X0,X1,X2,X3,X4) | ~v1_xboole_0(X4) [cnf transformation 93] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1268. k1_gate_1(k1_xboole_0) = k21_gate_1(X0,X1,X2,X3,X4) | ~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 94] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1269. k1_xboole_0 = k21_gate_1(X0,X1,X2,X3,X4) | v1_xboole_0(X0) [cnf transformation 94] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1270. k1_xboole_0 = k21_gate_1(X0,X1,X2,X3,X4) | v1_xboole_0(X1) [cnf transformation 94] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1271. k1_xboole_0 = k21_gate_1(X0,X1,X2,X3,X4) | v1_xboole_0(X2) [cnf transformation 94] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1272. k1_xboole_0 = k21_gate_1(X0,X1,X2,X3,X4) | v1_xboole_0(X3) [cnf transformation 94] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1273. k1_xboole_0 = k21_gate_1(X0,X1,X2,X3,X4) | v1_xboole_0(X4) [cnf transformation 94] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1274. k38_gate_1(X0,X1,X2,X3,X4,X5,X6) = k11_gate_1(X4,X5,k36_gate_1(X0,X1,X2,X3,X6)) [cnf transformation 82] {a:0,w:18,wCS:7200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1275. k37_gate_1(X0,X1,X2,X3,X4,X5,X6) = k10_gate_1(X4,X5,k36_gate_1(X0,X1,X2,X3,X6)) [cnf transformation 83] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1276. k40_gate_1(X0,X1,X2,X3,X4,X5,X6,X7,X8) = k11_gate_1(X6,X7,k38_gate_1(X0,X1,X2,X3,X4,X5,X8)) [cnf transformation 84] {a:0,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1277. k39_gate_1(X0,X1,X2,X3,X4,X5,X6,X7,X8) = k10_gate_1(X6,X7,k38_gate_1(X0,X1,X2,X3,X4,X5,X8)) [cnf transformation 85] {a:0,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1279. k6_gate_1(X0,X1) = k6_gate_1(X1,X0) [cnf transformation 6] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1280. k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1281. k4_gate_1(X0,X1) = k4_gate_1(X1,X0) [cnf transformation 5] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1282. k7_gate_1(X0,X1) = k7_gate_1(X1,X0) [cnf transformation 7] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1283. v1_xboole_0(k7_gate_1(X0,X1)) | v1_xboole_0(X0) [cnf transformation 95] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1284. ~v1_xboole_0(k6_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 96] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1285. v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 97] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1286. ~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 99] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1287. v1_xboole_0(k6_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 101] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1288. v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 103] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1289. ~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 105] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1290. ~v1_xboole_0(k7_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 107] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1291. v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 109] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1292. v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0) [cnf transformation 110] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1293. v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0) [cnf transformation 111] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1294. v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0) [cnf transformation 112] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1295. ~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 114] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1296. ~v1_xboole_0(k10_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 116] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1297. v1_xboole_0(k10_gate_1(X2,X0,X1)) | ~v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 118] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1298. v1_xboole_0(k10_gate_1(X0,X1,X2)) | ~v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 120] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1299. v1_xboole_0(k10_gate_1(X0,X2,X1)) | ~v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 122] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1300. ~v1_xboole_0(k11_gate_1(X2,X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 124] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1301. ~v1_xboole_0(k11_gate_1(X0,X1,X2)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 126] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1302. ~v1_xboole_0(k11_gate_1(X0,X2,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 128] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1303. ~v1_xboole_0(k10_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 130] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1304. ~v1_xboole_0(k10_gate_1(X0,X2,X1)) | v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 132] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1305. ~v1_xboole_0(k10_gate_1(X2,X0,X1)) | v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 134] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1306. v1_xboole_0(k10_gate_1(X0,X1,X2)) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 136] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1307. v1_xboole_0(k11_gate_1(X0,X1,X2)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 138] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1308. v1_xboole_0(k11_gate_1(X0,X2,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 140] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1309. v1_xboole_0(k11_gate_1(X2,X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 142] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1313. 1 <=> v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) [avatar definition]
1314. ~v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) <- (~1) [avatar component clause 1313] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1315. v1_xboole_0(k10_gate_1(sK1,sK2,sK0)) <- (1) [avatar component clause 1313] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1317. 2 <=> v1_xboole_0(sK37) [avatar definition]
1321. 3 <=> v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) [avatar definition]
1322. ~v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) <- (~3) [avatar component clause 1321] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1323. v1_xboole_0(k35_gate_1(sK1,sK2,sK3,sK4,sK0)) <- (3) [avatar component clause 1321] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1325. 4 <=> v1_xboole_0(sK38) [avatar definition]
1329. 5 <=> v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) [avatar definition]
1330. ~v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) <- (~5) [avatar component clause 1329] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1331. v1_xboole_0(k37_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) <- (5) [avatar component clause 1329] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1333. 6 <=> v1_xboole_0(sK39) [avatar definition]
1337. 7 <=> v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) [avatar definition]
1338. ~v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) <- (~7) [avatar component clause 1337] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1339. v1_xboole_0(k39_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) <- (7) [avatar component clause 1337] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1341. 8 <=> v1_xboole_0(sK40) [avatar definition]
1345. 9 <=> v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) [avatar definition]
1346. ~v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) <- (~9) [avatar component clause 1345] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1347. v1_xboole_0(k40_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK0)) <- (9) [avatar component clause 1345] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1349. 10 <=> v1_xboole_0(sK9) [avatar definition]
1352. 1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | 9 | 10 [avatar split clause 1236,1349,1345,1341,1337,1333,1329,1325,1321,1317,1313]
1353. ~2 | ~1 | 3 | 4 | 5 | 6 | 7 | 8 | 9 | 10 [avatar split clause 1233,1349,1345,1341,1337,1333,1329,1325,1321,1313,1317]
1354. 1 | 2 | ~4 | ~3 | 5 | 6 | 7 | 8 | 9 | 10 [avatar split clause 1224,1349,1345,1341,1337,1333,1329,1321,1325,1317,1313]
1355. ~2 | ~1 | ~4 | ~3 | 5 | 6 | 7 | 8 | 9 | 10 [avatar split clause 1221,1349,1345,1341,1337,1333,1329,1321,1325,1313,1317]
1356. 1 | 2 | 3 | 4 | ~6 | ~5 | 7 | 8 | 9 | 10 [avatar split clause 1188,1349,1345,1341,1337,1329,1333,1325,1321,1317,1313]
1357. ~2 | ~1 | 3 | 4 | ~6 | ~5 | 7 | 8 | 9 | 10 [avatar split clause 1185,1349,1345,1341,1337,1329,1333,1325,1321,1313,1317]
1358. 1 | 2 | ~4 | ~3 | ~6 | ~5 | 7 | 8 | 9 | 10 [avatar split clause 1176,1349,1345,1341,1337,1329,1333,1321,1325,1317,1313]
1359. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | 7 | 8 | 9 | 10 [avatar split clause 1173,1349,1345,1341,1337,1329,1333,1321,1325,1313,1317]
1360. 1 | 2 | 3 | 4 | 5 | 6 | ~8 | ~7 | 9 | 10 [avatar split clause 1044,1349,1345,1337,1341,1333,1329,1325,1321,1317,1313]
1361. ~2 | ~1 | 3 | 4 | 5 | 6 | ~8 | ~7 | 9 | 10 [avatar split clause 1041,1349,1345,1337,1341,1333,1329,1325,1321,1313,1317]
1362. 1 | 2 | ~4 | ~3 | 5 | 6 | ~8 | ~7 | 9 | 10 [avatar split clause 1032,1349,1345,1337,1341,1333,1329,1321,1325,1317,1313]
1363. ~2 | ~1 | ~4 | ~3 | 5 | 6 | ~8 | ~7 | 9 | 10 [avatar split clause 1029,1349,1345,1337,1341,1333,1329,1321,1325,1313,1317]
1364. 1 | 2 | 3 | 4 | ~6 | ~5 | ~8 | ~7 | 9 | 10 [avatar split clause 996,1349,1345,1337,1341,1329,1333,1325,1321,1317,1313]
1365. ~2 | ~1 | 3 | 4 | ~6 | ~5 | ~8 | ~7 | 9 | 10 [avatar split clause 993,1349,1345,1337,1341,1329,1333,1325,1321,1313,1317]
1366. 1 | 2 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | 9 | 10 [avatar split clause 984,1349,1345,1337,1341,1329,1333,1321,1325,1317,1313]
1367. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | 9 | 10 [avatar split clause 981,1349,1345,1337,1341,1329,1333,1321,1325,1313,1317]
1368. 1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | ~10 | ~9 [avatar split clause 468,1345,1349,1341,1337,1333,1329,1325,1321,1317,1313]
1369. ~2 | ~1 | 3 | 4 | 5 | 6 | 7 | 8 | ~10 | ~9 [avatar split clause 465,1345,1349,1341,1337,1333,1329,1325,1321,1313,1317]
1370. 1 | 2 | ~4 | ~3 | 5 | 6 | 7 | 8 | ~10 | ~9 [avatar split clause 456,1345,1349,1341,1337,1333,1329,1321,1325,1317,1313]
1371. ~2 | ~1 | ~4 | ~3 | 5 | 6 | 7 | 8 | ~10 | ~9 [avatar split clause 453,1345,1349,1341,1337,1333,1329,1321,1325,1313,1317]
1372. 1 | 2 | 3 | 4 | ~6 | ~5 | 7 | 8 | ~10 | ~9 [avatar split clause 420,1345,1349,1341,1337,1329,1333,1325,1321,1317,1313]
1373. ~2 | ~1 | 3 | 4 | ~6 | ~5 | 7 | 8 | ~10 | ~9 [avatar split clause 417,1345,1349,1341,1337,1329,1333,1325,1321,1313,1317]
1374. 1 | 2 | ~4 | ~3 | ~6 | ~5 | 7 | 8 | ~10 | ~9 [avatar split clause 408,1345,1349,1341,1337,1329,1333,1321,1325,1317,1313]
1375. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | 7 | 8 | ~10 | ~9 [avatar split clause 405,1345,1349,1341,1337,1329,1333,1321,1325,1313,1317]
1376. 1 | 2 | 3 | 4 | 5 | 6 | ~8 | ~7 | ~10 | ~9 [avatar split clause 276,1345,1349,1337,1341,1333,1329,1325,1321,1317,1313]
1377. ~2 | ~1 | 3 | 4 | 5 | 6 | ~8 | ~7 | ~10 | ~9 [avatar split clause 273,1345,1349,1337,1341,1333,1329,1325,1321,1313,1317]
1378. 1 | 2 | ~4 | ~3 | 5 | 6 | ~8 | ~7 | ~10 | ~9 [avatar split clause 264,1345,1349,1337,1341,1333,1329,1321,1325,1317,1313]
1379. ~2 | ~1 | ~4 | ~3 | 5 | 6 | ~8 | ~7 | ~10 | ~9 [avatar split clause 261,1345,1349,1337,1341,1333,1329,1321,1325,1313,1317]
1380. 1 | 2 | 3 | 4 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 [avatar split clause 228,1345,1349,1337,1341,1329,1333,1325,1321,1317,1313]
1381. ~2 | ~1 | 3 | 4 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 [avatar split clause 225,1345,1349,1337,1341,1329,1333,1325,1321,1313,1317]
1382. 1 | 2 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 [avatar split clause 216,1345,1349,1337,1341,1329,1333,1321,1325,1317,1313]
1383. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 [avatar split clause 213,1345,1349,1337,1341,1329,1333,1321,1325,1313,1317]
1385. 11 <=> v1_xboole_0(k21_gate_1(sK19,sK33,sK34,sK35,sK36)) [avatar definition]
1386. ~v1_xboole_0(k21_gate_1(sK19,sK33,sK34,sK35,sK36)) <- (~11) [avatar component clause 1385] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1387. v1_xboole_0(k21_gate_1(sK19,sK33,sK34,sK35,sK36)) <- (11) [avatar component clause 1385] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1388. 11 | ~10 [avatar split clause 212,1349,1385]
1389. 10 | ~11 [avatar split clause 211,1385,1349]
1391. 12 <=> v1_xboole_0(k18_gate_1(sK20,sK17,sK14,sK11,sK22)) [avatar definition]
1392. ~v1_xboole_0(k18_gate_1(sK20,sK17,sK14,sK11,sK22)) <- (~12) [avatar component clause 1391] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1393. v1_xboole_0(k18_gate_1(sK20,sK17,sK14,sK11,sK22)) <- (12) [avatar component clause 1391] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1395. 13 <=> v1_xboole_0(sK36) [avatar definition]
1398. 12 | ~13 [avatar split clause 210,1395,1391]
1399. 13 | ~12 [avatar split clause 209,1391,1395]
1401. 14 <=> v1_xboole_0(k14_gate_1(sK10,sK20,sK17,sK14)) [avatar definition]
1402. ~v1_xboole_0(k14_gate_1(sK10,sK20,sK17,sK14)) <- (~14) [avatar component clause 1401] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1403. v1_xboole_0(k14_gate_1(sK10,sK20,sK17,sK14)) <- (14) [avatar component clause 1401] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1405. 15 <=> v1_xboole_0(sK35) [avatar definition]
1406. v1_xboole_0(sK35) <- (15) [avatar component clause 1405] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1408. 14 | ~15 [avatar split clause 208,1405,1401]
1409. 15 | ~14 [avatar split clause 207,1401,1405]
1411. 16 <=> v1_xboole_0(k8_gate_1(sK13,sK20,sK17)) [avatar definition]
1412. ~v1_xboole_0(k8_gate_1(sK13,sK20,sK17)) <- (~16) [avatar component clause 1411] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1413. v1_xboole_0(k8_gate_1(sK13,sK20,sK17)) <- (16) [avatar component clause 1411] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1415. 17 <=> v1_xboole_0(sK34) [avatar definition]
1416. v1_xboole_0(sK34) <- (17) [avatar component clause 1415] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1418. 16 | ~17 [avatar split clause 206,1415,1411]
1419. 17 | ~16 [avatar split clause 205,1411,1415]
1421. 18 <=> v1_xboole_0(k2_gate_1(sK16,sK20)) [avatar definition]
1422. ~v1_xboole_0(k2_gate_1(sK16,sK20)) <- (~18) [avatar component clause 1421] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1423. v1_xboole_0(k2_gate_1(sK16,sK20)) <- (18) [avatar component clause 1421] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1425. 19 <=> v1_xboole_0(sK33) [avatar definition]
1426. v1_xboole_0(sK33) <- (19) [avatar component clause 1425] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1428. 18 | ~19 [avatar split clause 204,1425,1421]
1429. 19 | ~18 [avatar split clause 203,1421,1425]
1431. 20 <=> v1_xboole_0(k4_gate_1(sK32,sK21)) [avatar definition]
1432. ~v1_xboole_0(k4_gate_1(sK32,sK21)) <- (~20) [avatar component clause 1431] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1433. v1_xboole_0(k4_gate_1(sK32,sK21)) <- (20) [avatar component clause 1431] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1434. 20 | ~8 [avatar split clause 202,1341,1431]
1435. 8 | ~20 [avatar split clause 201,1431,1341]
1437. 21 <=> v1_xboole_0(k17_gate_1(sK29,sK30,sK31,sK16)) [avatar definition]
1438. ~v1_xboole_0(k17_gate_1(sK29,sK30,sK31,sK16)) <- (~21) [avatar component clause 1437] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1439. v1_xboole_0(k17_gate_1(sK29,sK30,sK31,sK16)) <- (21) [avatar component clause 1437] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1441. 22 <=> v1_xboole_0(sK32) [avatar definition]
1444. 21 | ~22 [avatar split clause 200,1441,1437]
1445. 22 | ~21 [avatar split clause 199,1437,1441]
1447. 23 <=> v1_xboole_0(k14_gate_1(sK17,sK14,sK11,sK22)) [avatar definition]
1448. ~v1_xboole_0(k14_gate_1(sK17,sK14,sK11,sK22)) <- (~23) [avatar component clause 1447] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1449. v1_xboole_0(k14_gate_1(sK17,sK14,sK11,sK22)) <- (23) [avatar component clause 1447] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1451. 24 <=> v1_xboole_0(sK31) [avatar definition]
1454. 23 | ~24 [avatar split clause 198,1451,1447]
1455. 24 | ~23 [avatar split clause 197,1447,1451]
1457. 25 <=> v1_xboole_0(k8_gate_1(sK10,sK17,sK14)) [avatar definition]
1458. ~v1_xboole_0(k8_gate_1(sK10,sK17,sK14)) <- (~25) [avatar component clause 1457] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1459. v1_xboole_0(k8_gate_1(sK10,sK17,sK14)) <- (25) [avatar component clause 1457] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1461. 26 <=> v1_xboole_0(sK30) [avatar definition]
1464. 25 | ~26 [avatar split clause 196,1461,1457]
1465. 26 | ~25 [avatar split clause 195,1457,1461]
1467. 27 <=> v1_xboole_0(k2_gate_1(sK13,sK17)) [avatar definition]
1468. ~v1_xboole_0(k2_gate_1(sK13,sK17)) <- (~27) [avatar component clause 1467] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1469. v1_xboole_0(k2_gate_1(sK13,sK17)) <- (27) [avatar component clause 1467] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1471. 28 <=> v1_xboole_0(sK29) [avatar definition]
1474. 27 | ~28 [avatar split clause 194,1471,1467]
1475. 28 | ~27 [avatar split clause 193,1467,1471]
1477. 29 <=> v1_xboole_0(k4_gate_1(sK28,sK18)) [avatar definition]
1478. ~v1_xboole_0(k4_gate_1(sK28,sK18)) <- (~29) [avatar component clause 1477] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1479. v1_xboole_0(k4_gate_1(sK28,sK18)) <- (29) [avatar component clause 1477] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1480. 29 | ~6 [avatar split clause 192,1333,1477]
1481. 6 | ~29 [avatar split clause 191,1477,1333]
1483. 30 <=> v1_xboole_0(k13_gate_1(sK26,sK27,sK13)) [avatar definition]
1484. ~v1_xboole_0(k13_gate_1(sK26,sK27,sK13)) <- (~30) [avatar component clause 1483] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1485. v1_xboole_0(k13_gate_1(sK26,sK27,sK13)) <- (30) [avatar component clause 1483] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1487. 31 <=> v1_xboole_0(sK28) [avatar definition]
1490. 30 | ~31 [avatar split clause 190,1487,1483]
1491. 31 | ~30 [avatar split clause 189,1483,1487]
1493. 32 <=> v1_xboole_0(k8_gate_1(sK14,sK11,sK22)) [avatar definition]
1494. ~v1_xboole_0(k8_gate_1(sK14,sK11,sK22)) <- (~32) [avatar component clause 1493] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1495. v1_xboole_0(k8_gate_1(sK14,sK11,sK22)) <- (32) [avatar component clause 1493] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1497. 33 <=> v1_xboole_0(sK27) [avatar definition]
1500. 32 | ~33 [avatar split clause 188,1497,1493]
1501. 33 | ~32 [avatar split clause 187,1493,1497]
1503. 34 <=> v1_xboole_0(k2_gate_1(sK10,sK14)) [avatar definition]
1504. ~v1_xboole_0(k2_gate_1(sK10,sK14)) <- (~34) [avatar component clause 1503] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1505. v1_xboole_0(k2_gate_1(sK10,sK14)) <- (34) [avatar component clause 1503] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1507. 35 <=> v1_xboole_0(sK26) [avatar definition]
1510. 34 | ~35 [avatar split clause 186,1507,1503]
1511. 35 | ~34 [avatar split clause 185,1503,1507]
1513. 36 <=> v1_xboole_0(k4_gate_1(sK25,sK15)) [avatar definition]
1514. ~v1_xboole_0(k4_gate_1(sK25,sK15)) <- (~36) [avatar component clause 1513] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1515. v1_xboole_0(k4_gate_1(sK25,sK15)) <- (36) [avatar component clause 1513] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1516. 36 | ~4 [avatar split clause 184,1325,1513]
1517. 4 | ~36 [avatar split clause 183,1513,1325]
1519. 37 <=> v1_xboole_0(k7_gate_1(sK24,sK10)) [avatar definition]
1520. ~v1_xboole_0(k7_gate_1(sK24,sK10)) <- (~37) [avatar component clause 1519] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1521. v1_xboole_0(k7_gate_1(sK24,sK10)) <- (37) [avatar component clause 1519] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1523. 38 <=> v1_xboole_0(sK25) [avatar definition]
1526. 37 | ~38 [avatar split clause 182,1523,1519]
1527. 38 | ~37 [avatar split clause 181,1519,1523]
1529. 39 <=> v1_xboole_0(k2_gate_1(sK22,sK11)) [avatar definition]
1530. ~v1_xboole_0(k2_gate_1(sK22,sK11)) <- (~39) [avatar component clause 1529] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1531. v1_xboole_0(k2_gate_1(sK22,sK11)) <- (39) [avatar component clause 1529] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1533. 40 <=> v1_xboole_0(sK24) [avatar definition]
1536. 39 | ~40 [avatar split clause 180,1533,1529]
1537. 40 | ~39 [avatar split clause 179,1529,1533]
1539. 41 <=> v1_xboole_0(k4_gate_1(sK23,sK12)) [avatar definition]
1540. ~v1_xboole_0(k4_gate_1(sK23,sK12)) <- (~41) [avatar component clause 1539] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1541. v1_xboole_0(k4_gate_1(sK23,sK12)) <- (41) [avatar component clause 1539] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1542. 41 | ~2 [avatar split clause 178,1317,1539]
1543. 2 | ~41 [avatar split clause 177,1539,1317]
1545. 42 <=> v1_xboole_0(k1_gate_1(sK22)) [avatar definition]
1546. ~v1_xboole_0(k1_gate_1(sK22)) <- (~42) [avatar component clause 1545] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1547. v1_xboole_0(k1_gate_1(sK22)) <- (42) [avatar component clause 1545] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1549. 43 <=> v1_xboole_0(sK23) [avatar definition]
1552. 42 | ~43 [avatar split clause 176,1549,1545]
1553. 43 | ~42 [avatar split clause 175,1545,1549]
1555. 44 <=> v1_xboole_0(k1_gate_1(sK0)) [avatar definition]
1556. ~v1_xboole_0(k1_gate_1(sK0)) <- (~44) [avatar component clause 1555] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1557. v1_xboole_0(k1_gate_1(sK0)) <- (44) [avatar component clause 1555] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1559. 45 <=> v1_xboole_0(sK22) [avatar definition]
1560. v1_xboole_0(sK22) <- (45) [avatar component clause 1559] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1562. 44 | ~45 [avatar split clause 174,1559,1555]
1563. 45 | ~44 [avatar split clause 173,1555,1559]
1565. 46 <=> v1_xboole_0(k34_gate_1(sK7,sK8)) [avatar definition]
1566. ~v1_xboole_0(k34_gate_1(sK7,sK8)) <- (~46) [avatar component clause 1565] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1567. v1_xboole_0(k34_gate_1(sK7,sK8)) <- (46) [avatar component clause 1565] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1569. 47 <=> v1_xboole_0(sK21) [avatar definition]
1572. 46 | ~47 [avatar split clause 172,1569,1565]
1573. 47 | ~46 [avatar split clause 171,1565,1569]
1575. 48 <=> v1_xboole_0(k6_gate_1(sK7,sK8)) [avatar definition]
1576. ~v1_xboole_0(k6_gate_1(sK7,sK8)) <- (~48) [avatar component clause 1575] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1577. v1_xboole_0(k6_gate_1(sK7,sK8)) <- (48) [avatar component clause 1575] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1579. 49 <=> v1_xboole_0(sK20) [avatar definition]
1580. v1_xboole_0(sK20) <- (49) [avatar component clause 1579] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1582. 48 | ~49 [avatar split clause 170,1579,1575]
1583. 49 | ~48 [avatar split clause 169,1575,1579]
1585. 50 <=> v1_xboole_0(k7_gate_1(sK7,sK8)) [avatar definition]
1586. ~v1_xboole_0(k7_gate_1(sK7,sK8)) <- (~50) [avatar component clause 1585] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1587. v1_xboole_0(k7_gate_1(sK7,sK8)) <- (50) [avatar component clause 1585] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1589. 51 <=> v1_xboole_0(sK19) [avatar definition]
1590. v1_xboole_0(sK19) <- (51) [avatar component clause 1589] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1591. ~v1_xboole_0(sK19) <- (~51) [avatar component clause 1589] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1592. 50 | ~51 [avatar split clause 168,1589,1585]
1593. 51 | ~50 [avatar split clause 167,1585,1589]
1595. 52 <=> v1_xboole_0(k34_gate_1(sK5,sK6)) [avatar definition]
1596. ~v1_xboole_0(k34_gate_1(sK5,sK6)) <- (~52) [avatar component clause 1595] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1597. v1_xboole_0(k34_gate_1(sK5,sK6)) <- (52) [avatar component clause 1595] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1599. 53 <=> v1_xboole_0(sK18) [avatar definition]
1602. 52 | ~53 [avatar split clause 166,1599,1595]
1603. 53 | ~52 [avatar split clause 165,1595,1599]
1605. 54 <=> v1_xboole_0(k6_gate_1(sK5,sK6)) [avatar definition]
1606. ~v1_xboole_0(k6_gate_1(sK5,sK6)) <- (~54) [avatar component clause 1605] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1607. v1_xboole_0(k6_gate_1(sK5,sK6)) <- (54) [avatar component clause 1605] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1609. 55 <=> v1_xboole_0(sK17) [avatar definition]
1610. v1_xboole_0(sK17) <- (55) [avatar component clause 1609] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1611. ~v1_xboole_0(sK17) <- (~55) [avatar component clause 1609] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1612. 54 | ~55 [avatar split clause 164,1609,1605]
1613. 55 | ~54 [avatar split clause 163,1605,1609]
1615. 56 <=> v1_xboole_0(k7_gate_1(sK5,sK6)) [avatar definition]
1616. ~v1_xboole_0(k7_gate_1(sK5,sK6)) <- (~56) [avatar component clause 1615] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1617. v1_xboole_0(k7_gate_1(sK5,sK6)) <- (56) [avatar component clause 1615] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1619. 57 <=> v1_xboole_0(sK16) [avatar definition]
1620. v1_xboole_0(sK16) <- (57) [avatar component clause 1619] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1622. 56 | ~57 [avatar split clause 162,1619,1615]
1623. 57 | ~56 [avatar split clause 161,1615,1619]
1625. 58 <=> v1_xboole_0(k34_gate_1(sK3,sK4)) [avatar definition]
1626. ~v1_xboole_0(k34_gate_1(sK3,sK4)) <- (~58) [avatar component clause 1625] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1627. v1_xboole_0(k34_gate_1(sK3,sK4)) <- (58) [avatar component clause 1625] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1629. 59 <=> v1_xboole_0(sK15) [avatar definition]
1632. 58 | ~59 [avatar split clause 160,1629,1625]
1633. 59 | ~58 [avatar split clause 159,1625,1629]
1635. 60 <=> v1_xboole_0(k6_gate_1(sK3,sK4)) [avatar definition]
1636. ~v1_xboole_0(k6_gate_1(sK3,sK4)) <- (~60) [avatar component clause 1635] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1637. v1_xboole_0(k6_gate_1(sK3,sK4)) <- (60) [avatar component clause 1635] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1639. 61 <=> v1_xboole_0(sK14) [avatar definition]
1640. v1_xboole_0(sK14) <- (61) [avatar component clause 1639] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1641. ~v1_xboole_0(sK14) <- (~61) [avatar component clause 1639] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1642. 60 | ~61 [avatar split clause 158,1639,1635]
1643. 61 | ~60 [avatar split clause 157,1635,1639]
1645. 62 <=> v1_xboole_0(k7_gate_1(sK3,sK4)) [avatar definition]
1646. ~v1_xboole_0(k7_gate_1(sK3,sK4)) <- (~62) [avatar component clause 1645] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1647. v1_xboole_0(k7_gate_1(sK3,sK4)) <- (62) [avatar component clause 1645] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1649. 63 <=> v1_xboole_0(sK13) [avatar definition]
1650. v1_xboole_0(sK13) <- (63) [avatar component clause 1649] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1651. ~v1_xboole_0(sK13) <- (~63) [avatar component clause 1649] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1652. 62 | ~63 [avatar split clause 156,1649,1645]
1653. 63 | ~62 [avatar split clause 155,1645,1649]
1655. 64 <=> v1_xboole_0(k34_gate_1(sK1,sK2)) [avatar definition]
1656. ~v1_xboole_0(k34_gate_1(sK1,sK2)) <- (~64) [avatar component clause 1655] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1657. v1_xboole_0(k34_gate_1(sK1,sK2)) <- (64) [avatar component clause 1655] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1659. 65 <=> v1_xboole_0(sK12) [avatar definition]
1662. 64 | ~65 [avatar split clause 154,1659,1655]
1663. 65 | ~64 [avatar split clause 153,1655,1659]
1665. 66 <=> v1_xboole_0(k6_gate_1(sK1,sK2)) [avatar definition]
1666. ~v1_xboole_0(k6_gate_1(sK1,sK2)) <- (~66) [avatar component clause 1665] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1667. v1_xboole_0(k6_gate_1(sK1,sK2)) <- (66) [avatar component clause 1665] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1669. 67 <=> v1_xboole_0(sK11) [avatar definition]
1670. v1_xboole_0(sK11) <- (67) [avatar component clause 1669] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1671. ~v1_xboole_0(sK11) <- (~67) [avatar component clause 1669] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1672. 66 | ~67 [avatar split clause 152,1669,1665]
1673. 67 | ~66 [avatar split clause 151,1665,1669]
1675. 68 <=> v1_xboole_0(k7_gate_1(sK1,sK2)) [avatar definition]
1676. ~v1_xboole_0(k7_gate_1(sK1,sK2)) <- (~68) [avatar component clause 1675] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1677. v1_xboole_0(k7_gate_1(sK1,sK2)) <- (68) [avatar component clause 1675] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1679. 69 <=> v1_xboole_0(sK10) [avatar definition]
1680. v1_xboole_0(sK10) <- (69) [avatar component clause 1679] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1681. ~v1_xboole_0(sK10) <- (~69) [avatar component clause 1679] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1682. 68 | ~69 [avatar split clause 150,1679,1675]
1683. 69 | ~68 [avatar split clause 149,1675,1679]
1685. 70 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
1687. v1_xboole_0(k1_xboole_0) <- (70) [avatar component clause 1685] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1688. 70 [avatar split clause 1237,1685]
1701. 73 <=> v1_xboole_0(sK0) [avatar definition]
1702. ~v1_xboole_0(sK0) <- (~73) [avatar component clause 1701] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1706. ~v1_xboole_0(sK22) <- (42) [resolution 1239,1547] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1710. $false <- (42, 45) [subsumption resolution 1706,1560] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1711. ~42 | ~45 [avatar contradiction clause 1710]
1713. v1_xboole_0(sK22) <- (~42) [resolution 1546,1238] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1714. v1_xboole_0(sK1) <- (~68) [resolution 1283,1676] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1715. v1_xboole_0(sK3) <- (~62) [resolution 1283,1646] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1716. v1_xboole_0(sK7) <- (~50) [resolution 1283,1586] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1717. v1_xboole_0(sK24) <- (~37) [resolution 1283,1520] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1719. 74 <=> v1_xboole_0(sK1) [avatar definition]
1721. v1_xboole_0(sK1) <- (74) [avatar component clause 1719] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1722. 74 | 68 [avatar split clause 1714,1675,1719]
1724. 75 <=> v1_xboole_0(sK3) [avatar definition]
1725. ~v1_xboole_0(sK3) <- (~75) [avatar component clause 1724] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1726. v1_xboole_0(sK3) <- (75) [avatar component clause 1724] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1727. 75 | 62 [avatar split clause 1715,1645,1724]
1729. 76 <=> v1_xboole_0(sK7) [avatar definition]
1730. ~v1_xboole_0(sK7) <- (~76) [avatar component clause 1729] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1732. 76 | 50 [avatar split clause 1716,1585,1729]
1733. ~v1_xboole_0(sK7) <- (48) [resolution 1284,1577] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1734. ~v1_xboole_0(sK3) <- (60) [resolution 1284,1637] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1735. ~76 | ~48 [avatar split clause 1733,1575,1729]
1736. ~75 | ~60 [avatar split clause 1734,1635,1724]
1739. ~v1_xboole_0(sK13) <- (~27) [resolution 1285,1468] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1740. ~v1_xboole_0(sK10) <- (~34) [resolution 1285,1504] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1743. $false <- (~27, 63) [subsumption resolution 1739,1650] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1744. 27 | ~63 [avatar contradiction clause 1743]
1745. $false <- (~34, 69) [subsumption resolution 1740,1680] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1746. 34 | ~69 [avatar contradiction clause 1745]
1747. ~v1_xboole_0(k6_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 1284,1279] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1754. 77 <=> v1_xboole_0(sK8) [avatar definition]
1755. v1_xboole_0(sK8) <- (77) [avatar component clause 1754] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1756. ~v1_xboole_0(sK8) <- (~77) [avatar component clause 1754] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1759. 78 <=> v1_xboole_0(sK4) [avatar definition]
1760. v1_xboole_0(sK4) <- (78) [avatar component clause 1759] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1761. ~v1_xboole_0(sK4) <- (~78) [avatar component clause 1759] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1764. v1_xboole_0(k2_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 1285,1280] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1767. 79 <=> v1_xboole_0(k2_gate_1(sK11,sK22)) [avatar definition]
1768. ~v1_xboole_0(k2_gate_1(sK11,sK22)) <- (~79) [avatar component clause 1767] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1769. v1_xboole_0(k2_gate_1(sK11,sK22)) <- (79) [avatar component clause 1767] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1771. 40 | 37 [avatar split clause 1717,1519,1533]
1780. 80 <=> v1_xboole_0(k4_gate_1(sK12,sK23)) [avatar definition]
1781. v1_xboole_0(k4_gate_1(sK12,sK23)) <- (80) [avatar component clause 1780] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1782. ~v1_xboole_0(k4_gate_1(sK12,sK23)) <- (~80) [avatar component clause 1780] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1785. 81 <=> v1_xboole_0(k4_gate_1(sK15,sK25)) [avatar definition]
1786. ~v1_xboole_0(k4_gate_1(sK15,sK25)) <- (~81) [avatar component clause 1785] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1787. v1_xboole_0(k4_gate_1(sK15,sK25)) <- (81) [avatar component clause 1785] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1790. 82 <=> v1_xboole_0(k4_gate_1(sK18,sK28)) [avatar definition]
1791. ~v1_xboole_0(k4_gate_1(sK18,sK28)) <- (~82) [avatar component clause 1790] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1792. v1_xboole_0(k4_gate_1(sK18,sK28)) <- (82) [avatar component clause 1790] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1795. 83 <=> v1_xboole_0(k4_gate_1(sK21,sK32)) [avatar definition]
1796. ~v1_xboole_0(k4_gate_1(sK21,sK32)) <- (~83) [avatar component clause 1795] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1797. v1_xboole_0(k4_gate_1(sK21,sK32)) <- (83) [avatar component clause 1795] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1806. v1_xboole_0(k7_gate_1(X1,X0)) | v1_xboole_0(X0) [superposition 1283,1282] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1809. 84 <=> v1_xboole_0(k7_gate_1(sK10,sK24)) [avatar definition]
1810. ~v1_xboole_0(k7_gate_1(sK10,sK24)) <- (~84) [avatar component clause 1809] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1811. v1_xboole_0(k7_gate_1(sK10,sK24)) <- (84) [avatar component clause 1809] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1817. v1_xboole_0(sK10) <- (~84) [resolution 1810,1283] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1819. ~v1_xboole_0(sK20) <- (~16) [resolution 1292,1412] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1823. ~49 | 16 [avatar split clause 1819,1411,1579]
1824. ~v1_xboole_0(sK17) <- (~16) [resolution 1293,1412] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1826. ~v1_xboole_0(sK13) <- (~16) [resolution 1294,1412] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1830. ~63 | 16 [avatar split clause 1826,1411,1649]
1850. ~v1_xboole_0(sK17) <- (~27) [resolution 1468,1764] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1854. v1_xboole_0(sK4) <- (~62) [resolution 1806,1646] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1860. 78 | 62 [avatar split clause 1854,1645,1759]
1876. 85 <=> v1_xboole_0(sK5) [avatar definition]
1877. ~v1_xboole_0(sK5) <- (~85) [avatar component clause 1876] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1878. v1_xboole_0(sK5) <- (85) [avatar component clause 1876] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1880. 86 <=> v1_xboole_0(sK6) [avatar definition]
1881. ~v1_xboole_0(sK6) <- (~86) [avatar component clause 1880] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1882. v1_xboole_0(sK6) <- (86) [avatar component clause 1880] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1885. 87 <=> v1_xboole_0(sK2) [avatar definition]
1886. ~v1_xboole_0(sK2) <- (~87) [avatar component clause 1885] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1887. v1_xboole_0(sK2) <- (87) [avatar component clause 1885] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1889. ~55 | 27 [avatar split clause 1850,1467,1609]
1890. ~55 | 16 [avatar split clause 1824,1411,1609]
1893. ~v1_xboole_0(sK5) <- (54) [resolution 1607,1284] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1895. ~85 | ~54 [avatar split clause 1893,1605,1876]
1927. ~v1_xboole_0(k4_gate_1(X1,X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [superposition 1289,1281] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1998. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK16) <- (~21) [superposition 1438,1259] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2027. 88 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
2028. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (88) [avatar component clause 2027] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
2029. ~v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (~88) [avatar component clause 2027] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
2052. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK17) <- (~12) [superposition 1392,1264] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2053. ~v1_xboole_0(sK17) <- (~12, 70) [subsumption resolution 2052,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2054. $false <- (~12, 55, 70) [subsumption resolution 2053,1610] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2055. 12 | ~55 | ~70 [avatar contradiction clause 2054]
2117. 69 | 84 [avatar split clause 1817,1809,1679]
2119. v1_xboole_0(sK14) | v1_xboole_0(sK10) <- (34) [resolution 1505,1286] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2123. ~v1_xboole_0(sK22) <- (~32) [resolution 1494,1293] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2126. $false <- (~32, 45) [subsumption resolution 2123,1560] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2127. 32 | ~45 [avatar contradiction clause 2126]
2140. ~v1_xboole_0(k36_gate_1(X2,X3,X0,X1,X4)) | v1_xboole_0(k11_gate_1(X2,X3,X4)) | v1_xboole_0(X1) [superposition 1300,1260] {a:1,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2141. ~v1_xboole_0(k36_gate_1(X7,X8,X5,X6,X9)) | v1_xboole_0(X6) | v1_xboole_0(X5) [superposition 1301,1260] {a:1,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2142. ~v1_xboole_0(k36_gate_1(X12,X13,X10,X11,X14)) | v1_xboole_0(k11_gate_1(X12,X13,X14)) | v1_xboole_0(X10) [superposition 1302,1260] {a:1,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2143. v1_xboole_0(k36_gate_1(X17,X18,X15,X16,X19)) | ~v1_xboole_0(X16) | ~v1_xboole_0(X15) [superposition 1307,1260] {a:1,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2144. v1_xboole_0(k36_gate_1(X22,X23,X20,X21,X24)) | ~v1_xboole_0(k11_gate_1(X22,X23,X24)) | ~v1_xboole_0(X20) [superposition 1308,1260] {a:1,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2145. v1_xboole_0(k36_gate_1(X27,X28,X25,X26,X29)) | ~v1_xboole_0(k11_gate_1(X27,X28,X29)) | ~v1_xboole_0(X26) [superposition 1309,1260] {a:1,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2147. ~v1_xboole_0(k35_gate_1(X2,X3,X0,X1,X4)) | v1_xboole_0(k11_gate_1(X2,X3,X4)) | v1_xboole_0(X1) | v1_xboole_0(X0) [superposition 1296,1261] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2148. v1_xboole_0(k35_gate_1(X7,X8,X5,X6,X9)) | ~v1_xboole_0(X5) | v1_xboole_0(k11_gate_1(X7,X8,X9)) | v1_xboole_0(X6) [superposition 1297,1261] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2149. v1_xboole_0(k35_gate_1(X12,X13,X10,X11,X14)) | ~v1_xboole_0(k11_gate_1(X12,X13,X14)) | v1_xboole_0(X11) | v1_xboole_0(X10) [superposition 1298,1261] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2150. v1_xboole_0(k35_gate_1(X17,X18,X15,X16,X19)) | ~v1_xboole_0(X16) | v1_xboole_0(k11_gate_1(X17,X18,X19)) | v1_xboole_0(X15) [superposition 1299,1261] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2151. ~v1_xboole_0(k35_gate_1(X22,X23,X20,X21,X24)) | v1_xboole_0(k11_gate_1(X22,X23,X24)) | ~v1_xboole_0(X21) | ~v1_xboole_0(X20) [superposition 1303,1261] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2152. ~v1_xboole_0(k35_gate_1(X27,X28,X25,X26,X29)) | v1_xboole_0(X26) | ~v1_xboole_0(k11_gate_1(X27,X28,X29)) | ~v1_xboole_0(X25) [superposition 1304,1261] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2153. ~v1_xboole_0(k35_gate_1(X32,X33,X30,X31,X34)) | v1_xboole_0(X30) | ~v1_xboole_0(k11_gate_1(X32,X33,X34)) | ~v1_xboole_0(X31) [superposition 1305,1261] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2154. v1_xboole_0(k35_gate_1(X37,X38,X35,X36,X39)) | ~v1_xboole_0(k11_gate_1(X37,X38,X39)) | ~v1_xboole_0(X36) | ~v1_xboole_0(X35) [superposition 1306,1261] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2177. ~v1_xboole_0(k38_gate_1(X20,X21,X22,X23,X18,X19,X24)) | v1_xboole_0(k36_gate_1(X20,X21,X22,X23,X24)) | v1_xboole_0(X19) [superposition 1300,1274] {a:1,w:18,wCS:7200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2178. ~v1_xboole_0(k38_gate_1(X27,X28,X29,X30,X25,X26,X31)) | v1_xboole_0(X26) | v1_xboole_0(X25) [superposition 1301,1274] {a:1,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2179. ~v1_xboole_0(k38_gate_1(X34,X35,X36,X37,X32,X33,X38)) | v1_xboole_0(k36_gate_1(X34,X35,X36,X37,X38)) | v1_xboole_0(X32) [superposition 1302,1274] {a:1,w:18,wCS:7200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2180. v1_xboole_0(k38_gate_1(X41,X42,X43,X44,X39,X40,X45)) | ~v1_xboole_0(X40) | ~v1_xboole_0(X39) [superposition 1307,1274] {a:1,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2181. v1_xboole_0(k38_gate_1(X48,X49,X50,X51,X46,X47,X52)) | ~v1_xboole_0(k36_gate_1(X48,X49,X50,X51,X52)) | ~v1_xboole_0(X46) [superposition 1308,1274] {a:1,w:18,wCS:7200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2182. v1_xboole_0(k38_gate_1(X55,X56,X57,X58,X53,X54,X59)) | ~v1_xboole_0(k36_gate_1(X55,X56,X57,X58,X59)) | ~v1_xboole_0(X54) [superposition 1309,1274] {a:1,w:18,wCS:7200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
2183. ~v1_xboole_0(k37_gate_1(X2,X3,X4,X5,X0,X1,X6)) | v1_xboole_0(k36_gate_1(X2,X3,X4,X5,X6)) | v1_xboole_0(X1) | v1_xboole_0(X0) [superposition 1296,1275] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2184. v1_xboole_0(k37_gate_1(X9,X10,X11,X12,X7,X8,X13)) | ~v1_xboole_0(X7) | v1_xboole_0(k36_gate_1(X9,X10,X11,X12,X13)) | v1_xboole_0(X8) [superposition 1297,1275] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2185. v1_xboole_0(k37_gate_1(X16,X17,X18,X19,X14,X15,X20)) | ~v1_xboole_0(k36_gate_1(X16,X17,X18,X19,X20)) | v1_xboole_0(X15) | v1_xboole_0(X14) [superposition 1298,1275] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2186. v1_xboole_0(k37_gate_1(X23,X24,X25,X26,X21,X22,X27)) | ~v1_xboole_0(X22) | v1_xboole_0(k36_gate_1(X23,X24,X25,X26,X27)) | v1_xboole_0(X21) [superposition 1299,1275] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2187. ~v1_xboole_0(k37_gate_1(X30,X31,X32,X33,X28,X29,X34)) | v1_xboole_0(k36_gate_1(X30,X31,X32,X33,X34)) | ~v1_xboole_0(X29) | ~v1_xboole_0(X28) [superposition 1303,1275] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2188. ~v1_xboole_0(k37_gate_1(X37,X38,X39,X40,X35,X36,X41)) | v1_xboole_0(X36) | ~v1_xboole_0(k36_gate_1(X37,X38,X39,X40,X41)) | ~v1_xboole_0(X35) [superposition 1304,1275] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2189. ~v1_xboole_0(k37_gate_1(X44,X45,X46,X47,X42,X43,X48)) | v1_xboole_0(X42) | ~v1_xboole_0(k36_gate_1(X44,X45,X46,X47,X48)) | ~v1_xboole_0(X43) [superposition 1305,1275] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2190. v1_xboole_0(k37_gate_1(X51,X52,X53,X54,X49,X50,X55)) | ~v1_xboole_0(k36_gate_1(X51,X52,X53,X54,X55)) | ~v1_xboole_0(X50) | ~v1_xboole_0(X49) [superposition 1306,1275] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2217. ~v1_xboole_0(k40_gate_1(X24,X25,X26,X27,X28,X29,X22,X23,X30)) | v1_xboole_0(k38_gate_1(X24,X25,X26,X27,X28,X29,X30)) | v1_xboole_0(X23) [superposition 1300,1276] {a:1,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2218. ~v1_xboole_0(k40_gate_1(X33,X34,X35,X36,X37,X38,X31,X32,X39)) | v1_xboole_0(X32) | v1_xboole_0(X31) [superposition 1301,1276] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2219. ~v1_xboole_0(k40_gate_1(X42,X43,X44,X45,X46,X47,X40,X41,X48)) | v1_xboole_0(k38_gate_1(X42,X43,X44,X45,X46,X47,X48)) | v1_xboole_0(X40) [superposition 1302,1276] {a:1,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2220. v1_xboole_0(k40_gate_1(X51,X52,X53,X54,X55,X56,X49,X50,X57)) | ~v1_xboole_0(X50) | ~v1_xboole_0(X49) [superposition 1307,1276] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2221. v1_xboole_0(k40_gate_1(X60,X61,X62,X63,X64,X65,X58,X59,X66)) | ~v1_xboole_0(k38_gate_1(X60,X61,X62,X63,X64,X65,X66)) | ~v1_xboole_0(X58) [superposition 1308,1276] {a:1,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2222. v1_xboole_0(k40_gate_1(X69,X70,X71,X72,X73,X74,X67,X68,X75)) | ~v1_xboole_0(k38_gate_1(X69,X70,X71,X72,X73,X74,X75)) | ~v1_xboole_0(X68) [superposition 1309,1276] {a:1,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2223. ~v1_xboole_0(k39_gate_1(X2,X3,X4,X5,X6,X7,X0,X1,X8)) | v1_xboole_0(k38_gate_1(X2,X3,X4,X5,X6,X7,X8)) | v1_xboole_0(X1) | v1_xboole_0(X0) [superposition 1296,1277] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2224. v1_xboole_0(k39_gate_1(X11,X12,X13,X14,X15,X16,X9,X10,X17)) | ~v1_xboole_0(X9) | v1_xboole_0(k38_gate_1(X11,X12,X13,X14,X15,X16,X17)) | v1_xboole_0(X10) [superposition 1297,1277] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2225. v1_xboole_0(k39_gate_1(X20,X21,X22,X23,X24,X25,X18,X19,X26)) | ~v1_xboole_0(k38_gate_1(X20,X21,X22,X23,X24,X25,X26)) | v1_xboole_0(X19) | v1_xboole_0(X18) [superposition 1298,1277] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2226. v1_xboole_0(k39_gate_1(X29,X30,X31,X32,X33,X34,X27,X28,X35)) | ~v1_xboole_0(X28) | v1_xboole_0(k38_gate_1(X29,X30,X31,X32,X33,X34,X35)) | v1_xboole_0(X27) [superposition 1299,1277] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2227. ~v1_xboole_0(k39_gate_1(X38,X39,X40,X41,X42,X43,X36,X37,X44)) | v1_xboole_0(k38_gate_1(X38,X39,X40,X41,X42,X43,X44)) | ~v1_xboole_0(X37) | ~v1_xboole_0(X36) [superposition 1303,1277] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2228. ~v1_xboole_0(k39_gate_1(X47,X48,X49,X50,X51,X52,X45,X46,X53)) | v1_xboole_0(X46) | ~v1_xboole_0(k38_gate_1(X47,X48,X49,X50,X51,X52,X53)) | ~v1_xboole_0(X45) [superposition 1304,1277] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2229. ~v1_xboole_0(k39_gate_1(X56,X57,X58,X59,X60,X61,X54,X55,X62)) | v1_xboole_0(X54) | ~v1_xboole_0(k38_gate_1(X56,X57,X58,X59,X60,X61,X62)) | ~v1_xboole_0(X55) [superposition 1305,1277] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2230. v1_xboole_0(k39_gate_1(X65,X66,X67,X68,X69,X70,X63,X64,X71)) | ~v1_xboole_0(k38_gate_1(X65,X66,X67,X68,X69,X70,X71)) | ~v1_xboole_0(X64) | ~v1_xboole_0(X63) [superposition 1306,1277] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2259. v1_xboole_0(sK3) | ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK4) <- (3) [resolution 2153,1323] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2265. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK4) <- (3, ~75) [subsumption resolution 2259,1725] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2266. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) <- (3, ~75, 78) [subsumption resolution 2265,1760] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2268. 89 <=> v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) [avatar definition]
2269. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) <- (89) [avatar component clause 2268] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2270. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) <- (~89) [avatar component clause 2268] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2271. ~89 | ~3 | 75 | ~78 [avatar split clause 2266,1759,1724,1321,2268]
2274. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) <- (~89) [resolution 2270,1307] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
2332. 90 <=> v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) [avatar definition]
2333. ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) <- (~90) [avatar component clause 2332] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2334. v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) <- (90) [avatar component clause 2332] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2378. 91 <=> v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) [avatar definition]
2379. ~v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) <- (~91) [avatar component clause 2378] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2380. v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) <- (91) [avatar component clause 2378] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2387. v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK5) <- (91) [resolution 2380,2179] {a:3,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2388. v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK6) <- (91) [resolution 2380,2177] {a:3,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2389. v1_xboole_0(sK6) | v1_xboole_0(sK5) <- (91) [resolution 2380,2178] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2393. v1_xboole_0(sK5) <- (~90, 91) [subsumption resolution 2387,2333] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2394. $false <- (~85, ~90, 91) [subsumption resolution 2393,1877] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2395. 85 | 90 | ~91 [avatar contradiction clause 2394]
2396. v1_xboole_0(sK6) <- (~90, 91) [subsumption resolution 2388,2333] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2397. $false <- (~86, ~90, 91) [subsumption resolution 2396,1881] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2398. 86 | 90 | ~91 [avatar contradiction clause 2397]
2399. v1_xboole_0(sK5) <- (~86, 91) [subsumption resolution 2389,1881] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2400. $false <- (~85, ~86, 91) [subsumption resolution 2399,1877] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2401. 85 | 86 | ~91 [avatar contradiction clause 2400]
2413. ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK6) <- (~91) [resolution 2379,2182] {a:3,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2414. ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK5) <- (~91) [resolution 2379,2181] {a:3,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2460. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | v1_xboole_0(sK3) <- (90) [resolution 2334,2142] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2464. v1_xboole_0(sK3) <- (~89, 90) [subsumption resolution 2460,2270] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2465. $false <- (~75, ~89, 90) [subsumption resolution 2464,1725] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2466. 75 | 89 | ~90 [avatar contradiction clause 2465]
2467. ~86 | ~90 | 91 [avatar split clause 2413,2378,2332,1880]
2507. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK2) | v1_xboole_0(sK1) <- (~64) [superposition 1656,1245] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2508. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) <- (~64) [superposition 1656,1244] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2522. v1_xboole_0(sK2) | v1_xboole_0(sK1) <- (~64, 70) [subsumption resolution 2507,1687] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2539. ~v1_xboole_0(sK1) <- (66) [resolution 1667,1284] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2541. $false <- (66, 74) [subsumption resolution 2539,1721] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2542. ~66 | ~74 [avatar contradiction clause 2541]
2543. v1_xboole_0(sK2) | v1_xboole_0(sK1) <- (~66) [resolution 1666,1287] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2544. v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK2) <- (1) [resolution 1315,1305] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2546. v1_xboole_0(sK0) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) <- (1) [resolution 1315,1303] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2563. v1_xboole_0(sK8) <- (~50) [resolution 1586,1806] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2566. $false <- (~50, ~77) [subsumption resolution 2563,1756] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2567. 50 | 77 [avatar contradiction clause 2566]
2574. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK17) <- (~14) [superposition 1402,1253] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2576. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK10) <- (~14) [superposition 1402,1251] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2577. ~v1_xboole_0(sK17) <- (~14, 70) [subsumption resolution 2574,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2578. $false <- (~14, 55, 70) [subsumption resolution 2577,1610] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2579. 14 | ~55 | ~70 [avatar contradiction clause 2578]
2580. ~v1_xboole_0(sK10) <- (~14, 70) [subsumption resolution 2576,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2581. $false <- (~14, 69, 70) [subsumption resolution 2580,1680] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2582. 14 | ~69 | ~70 [avatar contradiction clause 2581]
2587. ~v1_xboole_0(sK16) <- (~18) [resolution 1422,1285] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2591. ~57 | 18 [avatar split clause 2587,1421,1619]
2592. v1_xboole_0(sK16) <- (~21, 70) [subsumption resolution 1998,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2593. 57 | 21 | ~70 [avatar split clause 2592,1685,1437,1619]
2610. ~v1_xboole_0(sK4) | v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | v1_xboole_0(sK3) <- (~3) [resolution 1322,2150] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2614. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | v1_xboole_0(sK3) <- (~3, 78) [subsumption resolution 2610,1760] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2615. v1_xboole_0(sK3) <- (~3, 78, ~89) [subsumption resolution 2614,2270] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2616. $false <- (~3, ~75, 78, ~89) [subsumption resolution 2615,1725] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2617. 3 | 75 | ~78 | 89 [avatar contradiction clause 2616]
2634. ~v1_xboole_0(sK5) | v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK6) <- (~5) [resolution 1330,2184] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2638. ~v1_xboole_0(sK7) | v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK8) <- (~7) [resolution 1338,2224] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2640. ~v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK8) | v1_xboole_0(sK7) <- (~7) [resolution 2225,1338] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2642. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK4) <- (~90) [resolution 2333,2145] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2647. ~v1_xboole_0(sK8) | v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK7) <- (~7) [resolution 2226,1338] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2680. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK22) <- (~23) [superposition 1448,1254] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2683. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK17) <- (~23) [superposition 1448,1251] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2684. ~v1_xboole_0(sK22) <- (~23, 70) [subsumption resolution 2680,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2685. $false <- (~23, 45, 70) [subsumption resolution 2684,1560] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2686. 23 | ~45 | ~70 [avatar contradiction clause 2685]
2687. ~v1_xboole_0(sK17) <- (~23, 70) [subsumption resolution 2683,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2688. $false <- (~23, 55, 70) [subsumption resolution 2687,1610] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2689. 23 | ~55 | ~70 [avatar contradiction clause 2688]
2694. ~v1_xboole_0(sK10) <- (~25) [resolution 1458,1294] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2696. ~v1_xboole_0(sK17) <- (~25) [resolution 1458,1292] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2705. ~55 | 25 [avatar split clause 2696,1457,1609]
2706. ~69 | 25 [avatar split clause 2694,1457,1679]
2707. v1_xboole_0(sK10) <- (34, ~61) [subsumption resolution 2119,1641] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2708. 69 | ~34 | 61 [avatar split clause 2707,1639,1503,1679]
2709. ~v1_xboole_0(sK5) | v1_xboole_0(sK6) <- (~5, ~90) [subsumption resolution 2634,2333] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2710. ~v1_xboole_0(sK5) <- (~5, ~86, ~90) [subsumption resolution 2709,1881] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2711. ~85 | 5 | 86 | 90 [avatar split clause 2710,2332,1880,1329,1876]
2732. v1_xboole_0(sK2) <- (~68) [resolution 1676,1806] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2737. 87 | 68 [avatar split clause 2732,1675,1885]
2741. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) <- (~64, 70) [subsumption resolution 2508,1687] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2748. 45 | 42 [avatar split clause 1713,1545,1559]
2752. ~74 | ~87 | 73 | ~1 [avatar split clause 2546,1313,1701,1885,1719]
2756. ~74 | ~87 | 89 [avatar split clause 2274,2268,1885,1719]
2757. ~74 | ~87 | 64 | ~70 [avatar split clause 2741,1685,1655,1885,1719]
2758. 74 | 87 | 64 | ~70 [avatar split clause 2522,1685,1655,1885,1719]
2760. ~87 | ~73 | 74 | ~1 [avatar split clause 2544,1313,1719,1701,1885]
2761. 74 | 87 | 66 [avatar split clause 2543,1665,1885,1719]
2764. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) <- (78, ~90) [subsumption resolution 2642,1760] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2765. ~89 | ~78 | 90 [avatar split clause 2764,2332,1759,2268]
2768. ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) <- (85, ~91) [subsumption resolution 2414,1878] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2769. ~90 | ~85 | 91 [avatar split clause 2768,2378,1876,2332]
2780. v1_xboole_0(sK6) <- (~56) [resolution 1616,1806] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2783. $false <- (~56, ~86) [subsumption resolution 2780,1881] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2784. 56 | 86 [avatar contradiction clause 2783]
2788. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) <- (68) [resolution 1677,1290] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2790. ~v1_xboole_0(sK22) <- (~79) [resolution 1768,1764] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2793. $false <- (45, ~79) [subsumption resolution 2790,1560] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2794. ~45 | 79 [avatar contradiction clause 2793]
2833. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK27) <- (~30) [superposition 1484,1248] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2856. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK22) <- (~12) [superposition 1392,1267] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2860. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK20) <- (~12) [superposition 1392,1263] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2861. ~v1_xboole_0(sK22) <- (~12, 70) [subsumption resolution 2856,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2862. $false <- (~12, 45, 70) [subsumption resolution 2861,1560] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2863. 12 | ~45 | ~70 [avatar contradiction clause 2862]
2864. ~v1_xboole_0(sK20) <- (~12, 70) [subsumption resolution 2860,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2865. $false <- (~12, 49, 70) [subsumption resolution 2864,1580] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2866. 12 | ~49 | ~70 [avatar contradiction clause 2865]
2889. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | v1_xboole_0(sK4) | v1_xboole_0(sK3) <- (3) [resolution 1323,2147] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2895. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | v1_xboole_0(sK3) <- (3, ~78) [subsumption resolution 2889,1761] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2896. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) <- (3, ~75, ~78) [subsumption resolution 2895,1725] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2897. 89 | ~3 | 75 | 78 [avatar split clause 2896,1759,1724,1321,2268]
2902. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK33) <- (~11) [superposition 1386,1270] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2931. ~v1_xboole_0(sK2) <- (66) [resolution 1667,1747] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2934. $false <- (66, 87) [subsumption resolution 2931,1887] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2935. ~66 | ~87 [avatar contradiction clause 2934]
2956. v1_xboole_0(sK21) | ~v1_xboole_0(sK32) <- (83) [resolution 1797,1927] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2957. v1_xboole_0(sK32) | ~v1_xboole_0(sK21) <- (83) [resolution 1797,1289] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2974. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK14) <- (~23) [superposition 1448,1252] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2976. ~v1_xboole_0(sK14) <- (~23, 70) [subsumption resolution 2974,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2977. $false <- (~23, 61, 70) [subsumption resolution 2976,1640] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2978. 23 | ~61 | ~70 [avatar contradiction clause 2977]
2979. ~47 | 22 | ~83 [avatar split clause 2957,1795,1441,1569]
2980. ~22 | 47 | ~83 [avatar split clause 2956,1795,1569,1441]
3004. $false <- (~3, ~75, ~78, 89) [unit resulting resolution 1725,1761,2269,1322,2149] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3007. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | v1_xboole_0(sK4) | v1_xboole_0(sK3) <- (~3) [resolution 1322,2149] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3009. 3 | 75 | 78 | ~89 [avatar contradiction clause 3004]
3010. v1_xboole_0(sK4) | v1_xboole_0(sK3) <- (~3, 89) [subsumption resolution 3007,2269] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3011. v1_xboole_0(sK3) <- (~3, ~78, 89) [subsumption resolution 3010,1761] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3012. $false <- (~3, ~75, ~78, 89) [subsumption resolution 3011,1725] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3013. 3 | 75 | 78 | ~89 [avatar contradiction clause 3012]
3022. v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) <- (5) [resolution 1331,2187] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3023. v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK6) | v1_xboole_0(sK5) <- (5) [resolution 1331,2183] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3042. v1_xboole_0(sK33) <- (~11, 70) [subsumption resolution 2902,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3043. 19 | 11 | ~70 [avatar split clause 3042,1685,1385,1425]
3069. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | v1_xboole_0(sK4) <- (90) [resolution 2334,2140] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3070. v1_xboole_0(sK4) | v1_xboole_0(sK3) <- (90) [resolution 2334,2141] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3089. v1_xboole_0(sK4) <- (~75, 90) [subsumption resolution 3070,1725] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3090. 78 | 75 | ~90 [avatar split clause 3089,2332,1724,1759]
3093. v1_xboole_0(sK27) <- (~30, 70) [subsumption resolution 2833,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3094. 33 | 30 | ~70 [avatar split clause 3093,1685,1483,1497]
3160. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) <- (~3) [resolution 1322,2154] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3165. ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) <- (~3, 89) [subsumption resolution 3160,2269] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3166. ~v1_xboole_0(sK3) <- (~3, 78, 89) [subsumption resolution 3165,1760] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3167. $false <- (~3, 75, 78, 89) [subsumption resolution 3166,1726] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3168. 3 | ~75 | ~78 | ~89 [avatar contradiction clause 3167]
3212. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) <- (3) [resolution 1323,2151] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3214. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK3) <- (3, 78) [subsumption resolution 3212,1760] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3215. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) <- (3, 75, 78) [subsumption resolution 3214,1726] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3216. 89 | ~3 | ~75 | ~78 [avatar split clause 3215,1759,1724,1321,2268]
3219. 76 | 77 | ~91 | 7 [avatar split clause 2640,1337,2378,1754,1729]
3220. 76 | 91 | ~77 | 7 [avatar split clause 2647,1337,1754,2378,1729]
3221. 77 | 91 | ~76 | 7 [avatar split clause 2638,1337,1729,2378,1754]
3236. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK13) <- (~30) [superposition 1484,1249] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3239. v1_xboole_0(sK13) <- (~30, 70) [subsumption resolution 3236,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3240. $false <- (~30, ~63, 70) [subsumption resolution 3239,1651] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3241. 30 | 63 | ~70 [avatar contradiction clause 3240]
3253. 78 | 89 | ~90 [avatar split clause 3069,2332,2268,1759]
3319. ~v1_xboole_0(sK3) | v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | v1_xboole_0(sK4) <- (~3) [resolution 1322,2148] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3352. ~v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK7) <- (~7) [resolution 2230,1338] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3363. ~v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK7) <- (~7, 77) [subsumption resolution 3352,1755] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3415. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK3) <- (~90) [resolution 2333,2144] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3416. ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) <- (~90) [resolution 2333,2143] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3421. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) <- (75, ~90) [subsumption resolution 3415,1726] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3422. ~89 | ~75 | 90 [avatar split clause 3421,2332,1724,2268]
3423. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | v1_xboole_0(sK4) <- (~3, 75) [subsumption resolution 3319,1726] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3424. v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) <- (~3, 75, ~78) [subsumption resolution 3423,1761] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3425. 89 | 3 | ~75 | 78 [avatar split clause 3424,1759,1724,1321,2268]
3461. ~v1_xboole_0(sK0) | ~v1_xboole_0(sK2) <- (~89) [resolution 2270,1309] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
3462. ~v1_xboole_0(sK0) | ~v1_xboole_0(sK1) <- (~89) [resolution 2270,1308] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
3478. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK19) <- (~11) [superposition 1386,1269] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3479. v1_xboole_0(sK19) <- (~11, 70) [subsumption resolution 3478,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3480. $false <- (~11, ~51, 70) [subsumption resolution 3479,1591] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3481. 11 | 51 | ~70 [avatar contradiction clause 3480]
3492. ~87 | ~73 | 89 [avatar split clause 3461,2268,1701,1885]
3493. ~75 | ~78 | 90 [avatar split clause 3416,2332,1759,1724]
3501. v1_xboole_0(sK17) | v1_xboole_0(sK13) <- (27) [resolution 1469,1286] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3509. ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) <- (62) [resolution 1647,1290] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3518. ~v1_xboole_0(sK11) <- (~79) [resolution 1768,1285] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3520. $false <- (67, ~79) [subsumption resolution 3518,1670] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3521. ~67 | 79 [avatar contradiction clause 3520]
3565. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK11) <- (~23) [superposition 1448,1253] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3568. ~v1_xboole_0(sK11) <- (~23, 70) [subsumption resolution 3565,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3569. $false <- (~23, 67, 70) [subsumption resolution 3568,1670] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3570. 23 | ~67 | ~70 [avatar contradiction clause 3569]
3587. v1_xboole_0(sK5) <- (~56) [resolution 1616,1283] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3589. $false <- (~56, ~85) [subsumption resolution 3587,1877] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3590. 56 | 85 [avatar contradiction clause 3589]
3591. ~75 | ~78 | ~62 [avatar split clause 3509,1645,1759,1724]
3623. ~v1_xboole_0(sK11) <- (~32) [resolution 1494,1292] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3625. $false <- (~32, 67) [subsumption resolution 3623,1670] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3626. 32 | ~67 [avatar contradiction clause 3625]
3632. 63 | 55 | ~27 [avatar split clause 3501,1467,1609,1649]
3637. ~74 | ~73 | 89 [avatar split clause 3462,2268,1701,1719]
3649. ~76 | ~91 | 7 | ~77 [avatar split clause 3363,1754,1337,2378,1729]
3651. ~85 | ~86 | 90 | ~5 [avatar split clause 3022,1329,2332,1880,1876]
3652. 85 | 86 | 90 | ~5 [avatar split clause 3023,1329,2332,1880,1876]
3658. ~74 | ~87 | ~68 [avatar split clause 2788,1675,1885,1719]
3676. ~v1_xboole_0(sK8) | ~v1_xboole_0(sK7) <- (50) [resolution 1587,1290] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3690. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) <- (~58) [superposition 1626,1244] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3691. ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) <- (~58, 70) [subsumption resolution 3690,1687] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3692. ~v1_xboole_0(sK4) <- (~58, 70, 75) [subsumption resolution 3691,1726] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3693. ~78 | 58 | ~70 | ~75 [avatar split clause 3692,1724,1685,1625,1759]
3713. ~v1_xboole_0(sK0) | v1_xboole_0(sK2) | v1_xboole_0(sK1) <- (~1) [resolution 1314,1298] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3714. ~v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(sK2) <- (~1) [resolution 1314,1297] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3715. v1_xboole_0(sK17) | v1_xboole_0(sK20) | v1_xboole_0(sK13) <- (16) [resolution 1413,1295] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3716. ~v1_xboole_0(sK14) <- (~32) [resolution 1494,1294] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3744. ~61 | 32 [avatar split clause 3716,1493,1639]
3771. ~v1_xboole_0(sK14) <- (~34) [resolution 1504,1764] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3774. $false <- (~34, 61) [subsumption resolution 3771,1640] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3775. 34 | ~61 [avatar contradiction clause 3774]
3813. ~v1_xboole_0(sK14) <- (~25) [resolution 1458,1293] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3816. $false <- (~25, 61) [subsumption resolution 3813,1640] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3817. 25 | ~61 [avatar contradiction clause 3816]
3822. v1_xboole_0(sK14) | v1_xboole_0(sK17) | v1_xboole_0(sK10) <- (25) [resolution 1459,1295] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3826. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK26) <- (~30) [superposition 1484,1247] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3829. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK14) <- (~14) [superposition 1402,1254] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3831. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK20) <- (~14) [superposition 1402,1252] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3833. ~v1_xboole_0(sK14) <- (~14, 70) [subsumption resolution 3829,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3834. $false <- (~14, 61, 70) [subsumption resolution 3833,1640] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3835. 14 | ~61 | ~70 [avatar contradiction clause 3834]
3836. ~v1_xboole_0(sK20) <- (~14, 70) [subsumption resolution 3831,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3837. $false <- (~14, 49, 70) [subsumption resolution 3836,1580] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3838. 14 | ~49 | ~70 [avatar contradiction clause 3837]
3839. v1_xboole_0(sK26) <- (~30, 70) [subsumption resolution 3826,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3840. 35 | 30 | ~70 [avatar split clause 3839,1685,1483,1507]
3847. 69 | 55 | 61 | ~25 [avatar split clause 3822,1457,1639,1609,1679]
3852. ~v1_xboole_0(sK1) | v1_xboole_0(sK2) <- (~1, ~73) [subsumption resolution 3714,1702] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3853. 87 | ~74 | 1 | 73 [avatar split clause 3852,1701,1313,1719,1885]
3863. ~v1_xboole_0(sK20) <- (~18) [resolution 1422,1764] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3866. $false <- (~18, 49) [subsumption resolution 3863,1580] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3867. 18 | ~49 [avatar contradiction clause 3866]
3868. 74 | 87 | ~73 | 1 [avatar split clause 3713,1313,1701,1885,1719]
3870. 87 | 73 | ~74 | 1 [avatar split clause 3714,1313,1719,1701,1885]
3892. ~v1_xboole_0(sK6) <- (54) [resolution 1607,1747] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3894. ~86 | ~54 [avatar split clause 3892,1605,1880]
3895. ~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) <- (56) [resolution 1617,1290] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3929. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK14) | v1_xboole_0(sK17) | v1_xboole_0(sK20) | v1_xboole_0(sK10) <- (14) [superposition 1403,1250] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3942. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK22) | v1_xboole_0(sK11) | v1_xboole_0(sK14) | v1_xboole_0(sK17) <- (23) [superposition 1449,1250] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3963. v1_xboole_0(sK7) | ~v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK8) <- (7) [resolution 1339,2229] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3972. v1_xboole_0(sK22) | v1_xboole_0(sK11) | v1_xboole_0(sK14) | v1_xboole_0(sK17) <- (23, ~88) [subsumption resolution 3942,2029] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
3978. ~v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK8) <- (7, ~76) [subsumption resolution 3963,1730] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
3979. ~77 | ~91 | ~7 | 76 [avatar split clause 3978,1729,1337,2378,1754]
3997. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) <- (~52) [superposition 1596,1244] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4005. v1_xboole_0(sK15) | ~v1_xboole_0(sK25) <- (81) [resolution 1787,1927] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4006. v1_xboole_0(sK25) | ~v1_xboole_0(sK15) <- (81) [resolution 1787,1289] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4023. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK14) <- (~12) [superposition 1392,1265] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4026. ~v1_xboole_0(sK14) <- (~12, 70) [subsumption resolution 4023,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4027. $false <- (~12, 61, 70) [subsumption resolution 4026,1640] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4028. 12 | ~61 | ~70 [avatar contradiction clause 4027]
4029. ~59 | 38 | ~81 [avatar split clause 4006,1785,1523,1629]
4030. v1_xboole_0(sK22) | v1_xboole_0(sK11) | v1_xboole_0(sK14) <- (23, ~55, ~88) [subsumption resolution 3972,1611] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4031. 61 | 67 | 45 | ~23 | 55 | 88 [avatar split clause 4030,2027,1609,1447,1559,1669,1639]
4032. ~38 | 59 | ~81 [avatar split clause 4005,1785,1629,1523]
4037. ~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) <- (~52, 70) [subsumption resolution 3997,1687] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4038. ~85 | ~86 | 52 | ~70 [avatar split clause 4037,1685,1595,1880,1876]
4039. ~85 | ~86 | ~56 [avatar split clause 3895,1615,1880,1876]
4048. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK4) | v1_xboole_0(sK3) <- (~58) [superposition 1626,1245] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4050. v1_xboole_0(sK4) | v1_xboole_0(sK3) <- (~58, 70) [subsumption resolution 4048,1687] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4051. v1_xboole_0(sK3) <- (~58, 70, ~78) [subsumption resolution 4050,1761] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4052. 75 | 58 | ~70 | 78 [avatar split clause 4051,1759,1685,1625,1724]
4053. v1_xboole_0(sK4) | v1_xboole_0(sK3) <- (~60) [resolution 1636,1287] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4054. v1_xboole_0(sK3) <- (~60, ~78) [subsumption resolution 4053,1761] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4055. 75 | 60 | 78 [avatar split clause 4054,1759,1635,1724]
4066. v1_xboole_0(sK4) | ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK3) <- (3) [resolution 1323,2152] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4069. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) | ~v1_xboole_0(sK3) <- (3, ~78) [subsumption resolution 4066,1761] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4070. ~v1_xboole_0(k11_gate_1(sK1,sK2,sK0)) <- (3, 75, ~78) [subsumption resolution 4069,1726] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4071. ~89 | ~3 | ~75 | 78 [avatar split clause 4070,1759,1724,1321,2268]
4087. ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) <- (~5) [resolution 1330,2190] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4088. ~v1_xboole_0(sK6) | v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK5) <- (~5) [resolution 1330,2186] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4089. ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK6) | v1_xboole_0(sK5) <- (~5) [resolution 1330,2185] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4095. ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK5) <- (~5, 86) [subsumption resolution 4087,1882] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4096. ~85 | ~90 | 5 | ~86 [avatar split clause 4095,1880,1329,2332,1876]
4109. ~v1_xboole_0(sK28) | ~v1_xboole_0(sK18) <- (~82) [resolution 1791,1291] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4110. v1_xboole_0(sK28) | v1_xboole_0(sK18) <- (~82) [resolution 1791,1288] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4111. ~v1_xboole_0(sK32) | ~v1_xboole_0(sK21) <- (~83) [resolution 1796,1291] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4112. v1_xboole_0(sK32) | v1_xboole_0(sK21) <- (~83) [resolution 1796,1288] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4124. ~v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK8) <- (~9) [resolution 1346,2222] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4125. ~v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK7) <- (~9) [resolution 1346,2221] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4126. ~v1_xboole_0(sK8) | ~v1_xboole_0(sK7) <- (~9) [resolution 1346,2220] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4135. ~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) <- (~91) [resolution 2379,2180] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
4137. ~v1_xboole_0(sK5) <- (86, ~91) [subsumption resolution 4135,1882] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
4138. $false <- (85, 86, ~91) [subsumption resolution 4137,1878] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
4139. ~85 | ~86 | 91 [avatar contradiction clause 4138]
4140. 47 | 22 | 83 [avatar split clause 4112,1795,1441,1569]
4141. ~53 | ~31 | 82 [avatar split clause 4109,1790,1487,1599]
4142. ~47 | ~22 | 83 [avatar split clause 4111,1795,1441,1569]
4143. 53 | 31 | 82 [avatar split clause 4110,1790,1487,1599]
4144. v1_xboole_0(sK20) | v1_xboole_0(sK13) <- (16, ~55) [subsumption resolution 3715,1611] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4145. 63 | 49 | ~16 | 55 [avatar split clause 4144,1609,1411,1579,1649]
4146. ~77 | ~91 | 9 [avatar split clause 4124,1345,2378,1754]
4149. v1_xboole_0(sK20) | v1_xboole_0(sK16) <- (18) [resolution 1423,1286] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4169. ~v1_xboole_0(sK25) | ~v1_xboole_0(sK15) <- (~81) [resolution 1786,1291] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4170. v1_xboole_0(sK25) | v1_xboole_0(sK15) <- (~81) [resolution 1786,1288] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4177. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK31) <- (~21) [superposition 1438,1258] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4178. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK30) <- (~21) [superposition 1438,1257] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4179. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK29) <- (~21) [superposition 1438,1256] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4180. v1_xboole_0(sK29) <- (~21, 70) [subsumption resolution 4179,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4183. v1_xboole_0(sK31) <- (~21, 70) [subsumption resolution 4177,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4184. 24 | 21 | ~70 [avatar split clause 4183,1685,1437,1451]
4185. 28 | 21 | ~70 [avatar split clause 4180,1685,1437,1471]
4186. ~59 | ~38 | 81 [avatar split clause 4169,1785,1523,1629]
4187. 59 | 38 | 81 [avatar split clause 4170,1785,1523,1629]
4195. v1_xboole_0(sK30) <- (~21, 70) [subsumption resolution 4178,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4196. 26 | 21 | ~70 [avatar split clause 4195,1685,1437,1461]
4200. ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | v1_xboole_0(sK6) <- (~5, ~85) [subsumption resolution 4089,1877] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4201. 86 | ~90 | 5 | 85 [avatar split clause 4200,1876,1329,2332,1880]
4202. ~v1_xboole_0(sK6) | v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) <- (~5, ~85) [subsumption resolution 4088,1877] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4203. 90 | ~86 | 5 | 85 [avatar split clause 4202,1876,1329,1880,2332]
4206. ~v1_xboole_0(k7_gate_1(sK10,sK24)) <- (~37) [forward demodulation 1520,1282] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4207. ~84 | 37 [avatar split clause 4206,1519,1809]
4221. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK4) | v1_xboole_0(sK3) <- (58) [superposition 1627,1242] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4225. ~v1_xboole_0(sK4) | v1_xboole_0(sK3) <- (58, ~88) [subsumption resolution 4221,2029] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4226. ~v1_xboole_0(sK4) <- (58, ~75, ~88) [subsumption resolution 4225,1725] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4227. ~78 | ~58 | 75 | 88 [avatar split clause 4226,2027,1724,1625,1759]
4228. ~v1_xboole_0(sK4) <- (60) [resolution 1637,1747] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4230. ~78 | ~60 [avatar split clause 4228,1635,1759]
4238. v1_xboole_0(sK22) | v1_xboole_0(sK11) <- (79) [resolution 1769,1286] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4239. v1_xboole_0(sK12) | ~v1_xboole_0(sK23) <- (80) [resolution 1781,1927] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4240. v1_xboole_0(sK23) | ~v1_xboole_0(sK12) <- (80) [resolution 1781,1289] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4244. v1_xboole_0(sK2) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK1) <- (1) [resolution 1315,1304] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4246. v1_xboole_0(sK0) | v1_xboole_0(sK2) | v1_xboole_0(sK1) <- (1) [resolution 1315,1296] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4252. v1_xboole_0(sK22) | v1_xboole_0(sK11) | v1_xboole_0(sK14) <- (32) [resolution 1495,1295] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4255. v1_xboole_0(sK2) | v1_xboole_0(sK1) <- (89) [resolution 2269,1301] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
4256. v1_xboole_0(sK0) | v1_xboole_0(sK2) <- (89) [resolution 2269,1300] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
4261. ~43 | 65 | ~80 [avatar split clause 4239,1780,1659,1549]
4262. 67 | 45 | ~79 [avatar split clause 4238,1767,1559,1669]
4263. ~65 | 43 | ~80 [avatar split clause 4240,1780,1549,1659]
4264. v1_xboole_0(sK0) <- (~87, 89) [subsumption resolution 4256,1886] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
4265. 73 | 87 | ~89 [avatar split clause 4264,2268,1885,1701]
4266. ~v1_xboole_0(sK0) | ~v1_xboole_0(sK1) <- (1, ~87) [subsumption resolution 4244,1886] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4267. ~74 | ~73 | ~1 | 87 [avatar split clause 4266,1885,1313,1701,1719]
4268. v1_xboole_0(sK1) <- (~87, 89) [subsumption resolution 4255,1886] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
4269. 74 | 87 | ~89 [avatar split clause 4268,2268,1885,1719]
4271. v1_xboole_0(sK0) | v1_xboole_0(sK1) <- (1, ~87) [subsumption resolution 4246,1886] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4272. 74 | 73 | ~1 | 87 [avatar split clause 4271,1885,1313,1701,1719]
4273. 61 | 67 | 45 | ~32 [avatar split clause 4252,1493,1559,1669,1639]
4274. ~76 | ~91 | 9 [avatar split clause 4125,1345,2378,1729]
4277. v1_xboole_0(k4_gate_1(sK18,sK28)) <- (29) [forward demodulation 1479,1281] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4278. 82 | ~29 [avatar split clause 4277,1477,1790]
4282. v1_xboole_0(sK0) <- (~44) [resolution 1556,1238] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4283. 73 | 44 [avatar split clause 4282,1555,1701]
4287. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK7) <- (~46) [superposition 1566,1244] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4296. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK6) | v1_xboole_0(sK5) <- (~52) [superposition 1596,1245] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4298. v1_xboole_0(sK6) | v1_xboole_0(sK5) <- (~52, 70) [subsumption resolution 4296,1687] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4299. v1_xboole_0(sK5) <- (~52, 70, ~86) [subsumption resolution 4298,1881] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4300. 85 | 52 | ~70 | 86 [avatar split clause 4299,1880,1685,1595,1876]
4301. v1_xboole_0(sK6) | v1_xboole_0(sK5) <- (~54) [resolution 1606,1287] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4302. v1_xboole_0(sK5) <- (~54, ~86) [subsumption resolution 4301,1881] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4303. 85 | 54 | 86 [avatar split clause 4302,1880,1605,1876]
4309. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK2) | v1_xboole_0(sK1) <- (64) [superposition 1657,1242] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4310. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK1) | v1_xboole_0(sK2) <- (64) [superposition 1657,1241] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4317. v1_xboole_0(sK18) | ~v1_xboole_0(sK28) <- (82) [resolution 1792,1927] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4318. v1_xboole_0(sK28) | ~v1_xboole_0(sK18) <- (82) [resolution 1792,1289] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4333. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK36) | ~v1_xboole_0(sK35) | ~v1_xboole_0(sK34) | ~v1_xboole_0(sK33) | ~v1_xboole_0(sK19) <- (11) [superposition 1387,1268] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4341. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK11) <- (~12) [superposition 1392,1266] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4345. ~v1_xboole_0(sK11) <- (~12, 70) [subsumption resolution 4341,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4346. $false <- (~12, 67, 70) [subsumption resolution 4345,1670] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4347. 12 | ~67 | ~70 [avatar contradiction clause 4346]
4348. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK36) | ~v1_xboole_0(sK34) | ~v1_xboole_0(sK33) | ~v1_xboole_0(sK19) <- (11, 15) [subsumption resolution 4333,1406] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4349. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK36) | ~v1_xboole_0(sK33) | ~v1_xboole_0(sK19) <- (11, 15, 17) [subsumption resolution 4348,1416] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4350. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK36) | ~v1_xboole_0(sK19) <- (11, 15, 17, 19) [subsumption resolution 4349,1426] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4351. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK36) <- (11, 15, 17, 19, 51) [subsumption resolution 4350,1590] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4352. ~13 | 88 | ~11 | ~15 | ~17 | ~19 | ~51 [avatar split clause 4351,1589,1425,1415,1405,1385,2027,1395]
4353. ~31 | 53 | ~82 [avatar split clause 4317,1790,1599,1487]
4354. ~53 | 31 | ~82 [avatar split clause 4318,1790,1487,1599]
4356. ~v1_xboole_0(sK8) | ~v1_xboole_0(sK7) <- (~46, 70) [subsumption resolution 4287,1687] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4357. ~76 | ~77 | 46 | ~70 [avatar split clause 4356,1685,1565,1754,1729]
4359. ~76 | ~77 | ~50 [avatar split clause 3676,1585,1754,1729]
4360. ~76 | ~77 | 9 [avatar split clause 4126,1345,1754,1729]
4372. v1_xboole_0(k4_gate_1(sK21,sK32)) <- (20) [forward demodulation 1433,1281] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4373. 83 | ~20 [avatar split clause 4372,1431,1795]
4376. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK8) | v1_xboole_0(sK7) <- (46) [superposition 1567,1242] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4377. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK7) | v1_xboole_0(sK8) <- (46) [superposition 1567,1241] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4382. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK6) | v1_xboole_0(sK5) <- (52) [superposition 1597,1242] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4383. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK5) | v1_xboole_0(sK6) <- (52) [superposition 1597,1241] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4390. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK3) | v1_xboole_0(sK4) <- (58) [superposition 1627,1241] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4393. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK4) <- (58, 75) [subsumption resolution 4390,1726] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4394. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (58, 75, ~78) [subsumption resolution 4393,1761] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4395. 88 | ~58 | ~75 | 78 [avatar split clause 4394,1759,1724,1625,2027]
4397. ~v1_xboole_0(k1_xboole_0) <- (88) [resolution 2028,1239] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4399. $false <- (70, 88) [subsumption resolution 4397,1687] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4400. ~70 | ~88 [avatar contradiction clause 4399]
4401. 57 | 49 | ~18 [avatar split clause 4149,1421,1579,1619]
4402. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK17) | v1_xboole_0(sK20) | v1_xboole_0(sK10) <- (14, ~61) [subsumption resolution 3929,1641] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4405. 87 | ~74 | 88 | ~64 [avatar split clause 4310,1655,2027,1719,1885]
4407. 74 | ~87 | 88 | ~64 [avatar split clause 4309,1655,2027,1885,1719]
4408. v1_xboole_0(k4_gate_1(sK12,sK23)) <- (41) [forward demodulation 1541,1281] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4409. 80 | ~41 [avatar split clause 4408,1539,1780]
4412. v1_xboole_0(k7_gate_1(sK10,sK24)) <- (37) [forward demodulation 1521,1282] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4413. 84 | ~37 [avatar split clause 4412,1519,1809]
4418. ~v1_xboole_0(sK8) <- (48) [resolution 1577,1747] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4420. ~77 | ~48 [avatar split clause 4418,1575,1754]
4442. ~v1_xboole_0(sK24) | ~v1_xboole_0(sK10) <- (84) [resolution 1811,1290] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4450. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK27) | ~v1_xboole_0(sK26) <- (30) [superposition 1485,1246] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4462. v1_xboole_0(sK0) | v1_xboole_0(sK1) <- (89) [resolution 2269,1302] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
4480. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK22) | v1_xboole_0(sK11) | v1_xboole_0(sK14) | v1_xboole_0(sK17) | v1_xboole_0(sK20) <- (12) [superposition 1393,1262] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4489. v1_xboole_0(sK5) | ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK6) <- (5) [resolution 1331,2189] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4490. v1_xboole_0(sK6) | ~v1_xboole_0(k36_gate_1(sK1,sK2,sK3,sK4,sK0)) | ~v1_xboole_0(sK5) <- (5) [resolution 1331,2188] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4497. v1_xboole_0(sK8) | ~v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK7) <- (7) [resolution 1339,2228] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4498. v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK7) <- (7) [resolution 1339,2227] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4499. v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK8) | v1_xboole_0(sK7) <- (7) [resolution 1339,2223] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4501. v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK7) <- (9) [resolution 1347,2219] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4502. v1_xboole_0(k38_gate_1(sK1,sK2,sK3,sK4,sK5,sK6,sK0)) | v1_xboole_0(sK8) <- (9) [resolution 1347,2217] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4503. v1_xboole_0(sK8) | v1_xboole_0(sK7) <- (9) [resolution 1347,2218] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4508. v1_xboole_0(sK22) | v1_xboole_0(sK11) | v1_xboole_0(sK14) | v1_xboole_0(sK17) | v1_xboole_0(sK20) <- (12, ~88) [subsumption resolution 4480,2029] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4513. v1_xboole_0(sK17) | v1_xboole_0(sK20) | v1_xboole_0(sK10) <- (14, ~61, ~88) [subsumption resolution 4402,2029] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4514. v1_xboole_0(sK17) | v1_xboole_0(sK20) <- (14, ~61, ~69, ~88) [subsumption resolution 4513,1681] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4515. 49 | 55 | ~14 | 61 | 69 | 88 [avatar split clause 4514,2027,1679,1639,1401,1609,1579]
4516. ~76 | ~91 | 77 | ~7 [avatar split clause 4497,1337,1754,2378,1729]
4517. ~v1_xboole_0(sK7) | v1_xboole_0(sK8) <- (46, ~88) [subsumption resolution 4377,2029] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4518. 77 | ~76 | ~46 | 88 [avatar split clause 4517,2027,1565,1729,1754]
4519. ~v1_xboole_0(sK8) | v1_xboole_0(sK7) <- (46, ~88) [subsumption resolution 4376,2029] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4520. 76 | ~77 | ~46 | 88 [avatar split clause 4519,2027,1565,1754,1729]
4521. 76 | 77 | ~9 [avatar split clause 4503,1345,1754,1729]
4523. v1_xboole_0(sK5) | ~v1_xboole_0(sK6) <- (5, 90) [subsumption resolution 4489,2334] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4524. ~86 | 85 | ~5 | ~90 [avatar split clause 4523,2332,1329,1876,1880]
4525. ~v1_xboole_0(sK6) | v1_xboole_0(sK5) <- (52, ~88) [subsumption resolution 4382,2029] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4526. 85 | ~86 | ~52 | 88 [avatar split clause 4525,2027,1595,1880,1876]
4527. ~v1_xboole_0(sK5) | v1_xboole_0(sK6) <- (52, ~88) [subsumption resolution 4383,2029] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4528. 86 | ~85 | ~52 | 88 [avatar split clause 4527,2027,1595,1876,1880]
4529. v1_xboole_0(sK6) | ~v1_xboole_0(sK5) <- (5, 90) [subsumption resolution 4490,2334] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
4530. ~85 | 86 | ~5 | ~90 [avatar split clause 4529,2332,1329,1880,1876]
4531. 77 | 91 | ~9 [avatar split clause 4502,1345,2378,1754]
4532. 76 | 91 | ~9 [avatar split clause 4501,1345,2378,1729]
4533. 76 | 77 | 91 | ~7 [avatar split clause 4499,1337,2378,1754,1729]
4534. ~76 | ~77 | 91 | ~7 [avatar split clause 4498,1337,2378,1754,1729]
4535. ~v1_xboole_0(k4_gate_1(sK21,sK32)) <- (~20) [forward demodulation 1432,1281] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4536. ~83 | 20 [avatar split clause 4535,1431,1795]
4537. v1_xboole_0(sK22) | v1_xboole_0(sK14) | v1_xboole_0(sK17) | v1_xboole_0(sK20) <- (12, ~67, ~88) [subsumption resolution 4508,1671] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4538. v1_xboole_0(sK22) | v1_xboole_0(sK17) | v1_xboole_0(sK20) <- (12, ~61, ~67, ~88) [subsumption resolution 4537,1641] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4539. 49 | 55 | 45 | ~12 | 61 | 67 | 88 [avatar split clause 4538,2027,1669,1639,1391,1559,1609,1579]
4540. ~v1_xboole_0(k4_gate_1(sK12,sK23)) <- (~41) [forward demodulation 1540,1281] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4541. ~80 | 41 [avatar split clause 4540,1539,1780]
4542. v1_xboole_0(k2_gate_1(sK11,sK22)) <- (39) [forward demodulation 1531,1280] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4543. 79 | ~39 [avatar split clause 4542,1529,1767]
4544. v1_xboole_0(k4_gate_1(sK15,sK25)) <- (36) [forward demodulation 1515,1281] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4545. 81 | ~36 [avatar split clause 4544,1513,1785]
4546. ~v1_xboole_0(k4_gate_1(sK18,sK28)) <- (~29) [forward demodulation 1478,1281] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4547. ~82 | 29 [avatar split clause 4546,1477,1790]
4549. ~v1_xboole_0(sK0) <- (44) [resolution 1557,1239] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4550. ~73 | ~44 [avatar split clause 4549,1555,1701]
4555. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK8) | v1_xboole_0(sK7) <- (~46) [superposition 1566,1245] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4557. v1_xboole_0(sK8) | v1_xboole_0(sK7) <- (~46, 70) [subsumption resolution 4555,1687] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4558. v1_xboole_0(sK8) <- (~46, 70, ~76) [subsumption resolution 4557,1730] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4559. 77 | 46 | ~70 | 76 [avatar split clause 4558,1729,1685,1565,1754]
4560. v1_xboole_0(sK8) | v1_xboole_0(sK7) <- (~48) [resolution 1576,1287] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4561. v1_xboole_0(sK8) <- (~48, ~76) [subsumption resolution 4560,1730] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4562. 77 | 48 | 76 [avatar split clause 4561,1729,1575,1754]
4577. ~v1_xboole_0(sK23) | ~v1_xboole_0(sK12) <- (~80) [resolution 1782,1291] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4578. v1_xboole_0(sK23) | v1_xboole_0(sK12) <- (~80) [resolution 1782,1288] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4585. ~v1_xboole_0(sK0) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) <- (~1) [resolution 1314,1306] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4586. ~v1_xboole_0(sK2) | v1_xboole_0(sK0) | v1_xboole_0(sK1) <- (~1) [resolution 1314,1299] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4601. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK16) | ~v1_xboole_0(sK31) | ~v1_xboole_0(sK30) | ~v1_xboole_0(sK29) <- (21) [superposition 1439,1255] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4612. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK36) <- (~11) [superposition 1386,1273] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4613. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK35) <- (~11) [superposition 1386,1272] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4614. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK34) <- (~11) [superposition 1386,1271] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4617. v1_xboole_0(sK35) <- (~11, 70) [subsumption resolution 4613,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4620. v1_xboole_0(sK34) <- (~11, 70) [subsumption resolution 4614,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4623. v1_xboole_0(sK36) <- (~11, 70) [subsumption resolution 4612,1687] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4624. 13 | 11 | ~70 [avatar split clause 4623,1685,1385,1395]
4625. 15 | 11 | ~70 [avatar split clause 4617,1685,1385,1405]
4626. 17 | 11 | ~70 [avatar split clause 4620,1685,1385,1415]
4627. ~v1_xboole_0(sK16) | ~v1_xboole_0(sK31) | ~v1_xboole_0(sK30) | ~v1_xboole_0(sK29) <- (21, ~88) [subsumption resolution 4601,2029] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4628. ~v1_xboole_0(sK31) | ~v1_xboole_0(sK30) | ~v1_xboole_0(sK29) <- (21, 57, ~88) [subsumption resolution 4627,1620] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4629. ~28 | ~26 | ~24 | ~21 | ~57 | 88 [avatar split clause 4628,2027,1619,1437,1451,1461,1471]
4630. ~v1_xboole_0(sK13) | ~v1_xboole_0(sK27) | ~v1_xboole_0(sK26) <- (30, ~88) [subsumption resolution 4450,2029] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4631. ~35 | ~33 | ~63 | ~30 | 88 [avatar split clause 4630,2027,1483,1649,1497,1507]
4632. ~65 | ~43 | 80 [avatar split clause 4577,1780,1549,1659]
4633. 65 | 43 | 80 [avatar split clause 4578,1780,1549,1659]
4634. ~69 | ~40 | ~84 [avatar split clause 4442,1809,1533,1679]
4635. ~v1_xboole_0(sK0) | ~v1_xboole_0(sK1) <- (~1, 87) [subsumption resolution 4585,1887] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4636. ~74 | ~73 | 1 | ~87 [avatar split clause 4635,1885,1313,1701,1719]
4637. v1_xboole_0(sK0) | v1_xboole_0(sK1) <- (~1, 87) [subsumption resolution 4586,1887] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
4638. 74 | 73 | 1 | ~87 [avatar split clause 4637,1885,1313,1701,1719]
4639. 74 | 73 | ~89 [avatar split clause 4462,2268,1701,1719]
4640. ~v1_xboole_0(k2_gate_1(sK11,sK22)) <- (~39) [forward demodulation 1530,1280] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4641. ~79 | 39 [avatar split clause 4640,1529,1767]
4642. ~v1_xboole_0(k4_gate_1(sK15,sK25)) <- (~36) [forward demodulation 1514,1281] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
4643. ~81 | 36 [avatar split clause 4642,1513,1785]
4644. $false [avatar sat refutation 1352,1353,1354,1355,1356,1357,1358,1359,1360,1361,1362,1363,1364,1365,1366,1367,1368,1369,1370,1371,1372,1373,1374,1375,1376,1377,1378,1379,1380,1381,1382,1383,1388,1389,1398,1399,1408,1409,1418,1419,1428,1429,1434,1435,1444,1445,1454,1455,1464,1465,1474,1475,1480,1481,1490,1491,1500,1501,1510,1511,1516,1517,1526,1527,1536,1537,1542,1543,1552,1553,1562,1563,1572,1573,1582,1583,1592,1593,1602,1603,1612,1613,1622,1623,1632,1633,1642,1643,1652,1653,1662,1663,1672,1673,1682,1683,1688,1711,1722,1727,1732,1735,1736,1744,1746,1771,1823,1830,1860,1889,1890,1895,2055,2117,2127,2271,2395,2398,2401,2466,2467,2542,2567,2579,2582,2591,2593,2617,2686,2689,2705,2706,2708,2711,2737,2748,2752,2756,2757,2758,2760,2761,2765,2769,2784,2794,2863,2866,2897,2935,2978,2979,2980,3009,3013,3043,3090,3094,3168,3216,3219,3220,3221,3241,3253,3422,3425,3481,3492,3493,3521,3570,3590,3591,3626,3632,3637,3649,3651,3652,3658,3693,3744,3775,3817,3835,3838,3840,3847,3853,3867,3868,3870,3894,3979,4028,4029,4031,4032,4038,4039,4052,4055,4071,4096,4139,4140,4141,4142,4143,4145,4146,4184,4185,4186,4187,4196,4201,4203,4207,4227,4230,4261,4262,4263,4265,4267,4269,4272,4273,4274,4278,4283,4300,4303,4347,4352,4353,4354,4357,4359,4360,4373,4395,4400,4401,4405,4407,4409,4413,4420,4515,4516,4518,4520,4521,4524,4526,4528,4530,4531,4532,4533,4534,4536,4539,4541,4543,4545,4547,4550,4559,4562,4624,4625,4626,4629,4631,4632,4633,4634,4636,4638,4639,4641,4643] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:4080,thDist:-4080}
% SZS output end Proof for gate_1__t42_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 71
% Introduced skolems: 43
% 
% >>> Saturation
% Initial clauses: 1163
% Generated clauses: 5255
% Active clauses: 1190
% Passive clauses: 1500
% Final active clauses: 168
% Final passive clauses: 66
% Main loop iterations started: 1209
% 
% >>> Simplifying Inferences
% Duplicate literals: 108
% Fw subsumption resolutions: 409
% Fw demodulations: 170
% Bw demodulations: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 1158
% Forward subsumptions: 1304
% 
% >>> Generating Inferences
% Binary resolution: 783
% Unit resulting resolution: 53
% Forward superposition: 634
% Backward superposition: 182
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 266
% Split components: 821
% Sat splitting refutations: 132
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 671
% SAT solver unit clauses: 3
% SAT solver binary clauses: 271
% 
% Memory used [KB]: 11129
% Time elapsed: 0.191 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% property evaluation: 0.001 s
% SAT solver time: 0.002 s
% minimizing solver time: 0.002 s
% sat proof minimization: 0.001 s
% term sharing: 0.048 s
% splitting model update: 0.001 s ( own 0.0000 s ) 
% congruence closure: 0.005 s ( own 0.003 s ) 
% unit clause index maintenance: 0.001 s
% backward superposition index maintenance: 0.001 s
% forward subsumption: 0.010 s
% forward subsumption resolution: 0.004 s
% backward subsumption: 0.001 s
% forward demodulation: 0.001 s
% resolution: 0.002 s
% unit resulting resolution: 0.141 s ( own 0.095 s ) 
% superposition: 0.004 s ( own 0.003 s ) 
% passive container maintenance: 0.001 s
% other: 0.191 s ( own 0.014 s ) 

