i: [1,62,0,2,1,1,0,0] t31_funct_1
a: [3,99,0,2,1,62]
s: 99
i: [1,60,0,2,1,1,0,0] t31_funct_1
a: [3,109,0,2,1,60]
s: 109
i: [1,59,0,2,1,1,0,0] t31_funct_1
a: [3,114,0,2,1,59]
s: 114
i: [1,58,0,2,1,1,0,0] t31_funct_1
a: [3,119,0,2,1,58]
s: 119
i: [1,57,0,2,1,1,0,0] t31_funct_1
a: [3,124,0,2,1,57]
s: 124
i: [1,61,0,7,1,1,0,0] t31_funct_1
a: [3,104,0,7,1,61]
s: 104
i: [1,95,0,2,1,0,0,1] rc1_funct_1
a: [3,129,0,2,1,95]
s: 129
i: [1,94,0,2,1,0,0,1] rc1_funct_1
a: [3,134,0,2,1,94]
s: 134
i: [1,63,0,3,1,0,0,1] dt_k4_relat_1
s: 63
i: [1,66,0,3,1,0,0,1] fc4_funct_1
s: 66
i: [1,67,0,3,1,0,0,1] fc4_funct_1
s: 67
i: [1,68,0,3,1,0,0,1] fc3_funct_1
s: 68
i: [1,69,0,3,1,0,0,1] fc3_funct_1
s: 69
i: [1,70,0,3,1,0,0,4] fc28_relat_1
s: 70
i: [1,81,0,3,1,0,0,2] reflexivity_r1_tarski
s: 81
i: [1,71,0,4,1,0,0,4] fc28_relat_1
s: 71
i: [1,72,0,4,1,0,0,4] fc28_relat_1
s: 72
i: [1,80,0,4,1,0,0,3] existence_m1_subset_1
s: 80
i: [1,82,0,4,1,0,0,1] dt_k3_relat_1
s: 82
i: [1,87,0,4,1,0,0,4] rc3_relat_1
s: 87
i: [1,64,0,5,1,0,0,2] rd2_relat_1
s: 64
i: [1,65,0,5,1,0,0,1] rd1_relat_1
s: 65
i: [1,73,0,5,1,0,0,2] t45_relat_1
s: 73
i: [1,74,0,5,1,0,0,2] t45_relat_1
s: 74
i: [1,88,0,5,1,0,0,4] rc3_relat_1
s: 88
i: [1,89,0,5,1,0,0,4] rc3_relat_1
s: 89
i: [1,78,0,7,2,0,0,2] t3_subset
s: 78
i: [1,79,0,7,2,0,0,2] t3_subset
s: 79
d: [2,137,1,5,1,0,40,78,80]
s: 137
i: [1,75,0,8,3,0,0,3] cc2_relat_1
s: 75
i: [1,90,0,11,4,0,0,4] fc29_relat_1
s: 90
d: [2,139,1,6,2,0,40,75,80]
s: 139
i: [1,92,0,11,4,0,0,4] fc30_relat_1
s: 92
d: [2,140,1,7,3,0,40,75,79]
s: 140
i: [1,83,0,12,5,0,0,1] fc2_funct_1
s: 83
i: [1,84,0,12,5,0,0,1] fc2_funct_1
s: 84
i: [1,85,0,12,4,0,0,3] cc5_relat_1
s: 85
i: [1,86,0,12,4,0,0,3] cc6_relat_1
s: 86
d: [2,144,1,10,3,0,40,85,80]
s: 144
i: [1,91,0,12,4,0,0,4] fc29_relat_1
s: 91
i: [1,93,0,12,4,0,0,4] fc30_relat_1
s: 93
d: [2,148,1,12,4,1,44,91,104]
d: [2,149,1,10,3,2,21,148,124]
d: [2,150,1,8,2,3,21,149,114]
s: 150
i: [1,77,0,19,7,0,0,1] t25_funct_1
s: 77
d: [2,151,1,12,4,1,44,93,104]
d: [2,152,1,10,3,2,21,151,114]
d: [2,153,1,8,2,3,21,152,124]
s: 153
i: [1,76,0,20,6,0,0,2] t15_funct_1
s: 76
d: [2,158,1,20,6,1,44,76,104]
d: [2,159,1,13,5,1,21,158,65]
d: [2,160,1,11,4,2,21,159,119]
d: [2,161,1,9,3,3,21,160,124]
d: [2,162,1,7,2,4,21,161,109]
d: [2,163,1,5,1,5,21,162,114]
a: [3,167,1,5,1,163]
s: 167
d: [2,169,2,14,6,1,40,167,77]
d: [2,171,2,12,5,2,21,169,99]
d: [2,172,2,10,4,3,21,171,119]
d: [2,173,2,8,3,4,21,172,124]
d: [2,174,2,6,2,5,21,173,109]
d: [2,175,2,4,1,6,21,174,114]
d: [2,176,2,4,1,7,22,175,104]
d: [2,177,2,0,0,7,21,176,67]
e: 177
f: 177,163,57,58,59,60,61,62
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_1__t31_funct_1
% SZS output start Proof for funct_1__t31_funct_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & v1_funct_1(X1) & v1_relat_1(X1)) => v2_funct_1(X0))) [input t31_funct_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & v1_funct_1(X1) & v1_relat_1(X1)) => v2_funct_1(X0))) [negated conjecture 1]
18. ! [X0] : (v2_funct_1(k4_relat_1(X0)) & v1_relat_1(k4_relat_1(X0))) [input fc4_funct_1]
21. ! [X0] : k9_xtuple_0(k4_relat_1(X0)) = X0 [input rd1_relat_1]
24. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (k9_xtuple_0(k3_relat_1(X1,X0)) = k9_xtuple_0(X1) => r1_tarski(k10_xtuple_0(X1),k9_xtuple_0(X0))))) [input t15_funct_1]
25. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ((r1_tarski(k10_xtuple_0(X1),k9_xtuple_0(X0)) & v2_funct_1(k3_relat_1(X1,X0))) => v2_funct_1(X1)))) [input t25_funct_1]
30. ? [X0] : ((~v2_funct_1(X0) & ? [X1] : (k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
31. ? [X0] : (~v2_funct_1(X0) & ? [X1] : (k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 30]
33. ! [X0] : (! [X1] : ((r1_tarski(k10_xtuple_0(X1),k9_xtuple_0(X0)) | k9_xtuple_0(k3_relat_1(X1,X0)) != k9_xtuple_0(X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 24]
34. ! [X0] : (! [X1] : (r1_tarski(k10_xtuple_0(X1),k9_xtuple_0(X0)) | k9_xtuple_0(k3_relat_1(X1,X0)) != k9_xtuple_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 33]
35. ! [X0] : (! [X1] : ((v2_funct_1(X1) | (~r1_tarski(k10_xtuple_0(X1),k9_xtuple_0(X0)) | ~v2_funct_1(k3_relat_1(X1,X0)))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 25]
36. ! [X0] : (! [X1] : (v2_funct_1(X1) | ~r1_tarski(k10_xtuple_0(X1),k9_xtuple_0(X0)) | ~v2_funct_1(k3_relat_1(X1,X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 35]
47. ? [X0] : (~v2_funct_1(X0) & ? [X1] : (k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (~v2_funct_1(sK0) & ? [X1] : (k3_relat_1(sK0,X1) = k4_relat_1(k9_xtuple_0(sK0)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
48. ? [X1] : (k3_relat_1(sK0,X1) = k4_relat_1(k9_xtuple_0(sK0)) & v1_funct_1(X1) & v1_relat_1(X1)) => (k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
49. ~v2_funct_1(sK0) & (k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 31,48,47]
57. v1_relat_1(sK0) [cnf transformation 49] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. v1_funct_1(sK0) [cnf transformation 49] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
59. v1_relat_1(sK1) [cnf transformation 49] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
60. v1_funct_1(sK1) [cnf transformation 49] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
61. k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) [cnf transformation 49] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
62. ~v2_funct_1(sK0) [cnf transformation 49] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. k9_xtuple_0(k4_relat_1(X0)) = X0 [cnf transformation 21] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. v2_funct_1(k4_relat_1(X0)) [cnf transformation 18] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
76. k9_xtuple_0(k3_relat_1(X1,X0)) != k9_xtuple_0(X1) | r1_tarski(k10_xtuple_0(X1),k9_xtuple_0(X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 34] {a:0,w:20,wCS:8000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
77. ~r1_tarski(k10_xtuple_0(X1),k9_xtuple_0(X0)) | v2_funct_1(X1) | ~v2_funct_1(k3_relat_1(X1,X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
97. 1 <=> v2_funct_1(sK0) [avatar definition]
99. ~v2_funct_1(sK0) <- (~1) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. ~1 [avatar split clause 62,97]
102. 2 <=> k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) [avatar definition]
104. k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) <- (2) [avatar component clause 102] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 2 [avatar split clause 61,102]
107. 3 <=> v1_funct_1(sK1) [avatar definition]
109. v1_funct_1(sK1) <- (3) [avatar component clause 107] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. 3 [avatar split clause 60,107]
112. 4 <=> v1_relat_1(sK1) [avatar definition]
114. v1_relat_1(sK1) <- (4) [avatar component clause 112] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
115. 4 [avatar split clause 59,112]
117. 5 <=> v1_funct_1(sK0) [avatar definition]
119. v1_funct_1(sK0) <- (5) [avatar component clause 117] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
120. 5 [avatar split clause 58,117]
122. 6 <=> v1_relat_1(sK0) [avatar definition]
124. v1_relat_1(sK0) <- (6) [avatar component clause 122] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
125. 6 [avatar split clause 57,122]
158. k9_xtuple_0(sK0) != k9_xtuple_0(k4_relat_1(k9_xtuple_0(sK0))) | r1_tarski(k10_xtuple_0(sK0),k9_xtuple_0(sK1)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2) [superposition 76,104] {a:1,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
159. r1_tarski(k10_xtuple_0(sK0),k9_xtuple_0(sK1)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2) [subsumption resolution 158,65] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
160. r1_tarski(k10_xtuple_0(sK0),k9_xtuple_0(sK1)) | ~v1_relat_1(sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, 5) [subsumption resolution 159,119] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
161. r1_tarski(k10_xtuple_0(sK0),k9_xtuple_0(sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, 5, 6) [subsumption resolution 160,124] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
162. r1_tarski(k10_xtuple_0(sK0),k9_xtuple_0(sK1)) | ~v1_relat_1(sK1) <- (2, 3, 5, 6) [subsumption resolution 161,109] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
163. r1_tarski(k10_xtuple_0(sK0),k9_xtuple_0(sK1)) <- (2, 3, 4, 5, 6) [subsumption resolution 162,114] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
165. 9 <=> r1_tarski(k10_xtuple_0(sK0),k9_xtuple_0(sK1)) [avatar definition]
167. r1_tarski(k10_xtuple_0(sK0),k9_xtuple_0(sK1)) <- (9) [avatar component clause 165] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
168. 9 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar split clause 163,122,117,112,107,102,165]
169. v2_funct_1(sK0) | ~v2_funct_1(k3_relat_1(sK0,sK1)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (9) [resolution 167,77] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
171. ~v2_funct_1(k3_relat_1(sK0,sK1)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1, 9) [subsumption resolution 169,99] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
172. ~v2_funct_1(k3_relat_1(sK0,sK1)) | ~v1_relat_1(sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1, 5, 9) [subsumption resolution 171,119] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
173. ~v2_funct_1(k3_relat_1(sK0,sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1, 5, 6, 9) [subsumption resolution 172,124] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
174. ~v2_funct_1(k3_relat_1(sK0,sK1)) | ~v1_relat_1(sK1) <- (~1, 3, 5, 6, 9) [subsumption resolution 173,109] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
175. ~v2_funct_1(k3_relat_1(sK0,sK1)) <- (~1, 3, 4, 5, 6, 9) [subsumption resolution 174,114] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
176. ~v2_funct_1(k4_relat_1(k9_xtuple_0(sK0))) <- (~1, 2, 3, 4, 5, 6, 9) [forward demodulation 175,104] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
177. $false <- (~1, 2, 3, 4, 5, 6, 9) [subsumption resolution 176,67] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
178. 1 | ~2 | ~3 | ~4 | ~5 | ~6 | ~9 [avatar contradiction clause 177]
179. $false [avatar sat refutation 100,105,110,115,120,125,168,178] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:45,thDist:-45}
% SZS output end Proof for funct_1__t31_funct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 39
% Generated clauses: 85
% Active clauses: 38
% Passive clauses: 50
% Final active clauses: 38
% Final passive clauses: 4
% Main loop iterations started: 46
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 17
% Fw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 12
% Forward superposition: 7
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 8
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

