i: [1,99,0,9,2,1,0,0] t19_subset_1
a: [3,146,0,4,1,99]
s: 146
i: [1,97,0,4,1,1,0,0] t19_subset_1
a: [3,153,0,4,1,97]
s: 153
a: [3,143,0,5,1,99]
s: 143
i: [1,100,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,158,0,2,1,100]
s: 158
i: [1,101,0,2,1,0,0,2] fc1_xboole_0
a: [3,163,0,2,1,101]
s: 163
i: [1,135,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,173,0,2,1,135]
s: 173
i: [1,136,0,2,1,0,0,2] rc2_xboole_0
a: [3,178,0,2,1,136]
s: 178
i: [1,137,0,2,1,0,0,2] rc1_xboole_0
a: [3,183,0,2,1,137]
s: 183
i: [1,139,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,188,0,2,1,139]
s: 188
i: [1,138,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,193,0,2,1,138]
s: 193
i: [1,103,0,3,1,0,0,1] fc1_subset_1
s: 103
i: [1,104,0,3,1,0,0,1] t2_xboole_1
s: 104
i: [1,123,0,3,1,0,0,1] rc2_subset_1
s: 123
i: [1,124,0,3,1,0,0,1] reflexivity_r1_tarski
s: 124
i: [1,102,0,3,1,0,0,255] d2_xboole_0
a: [3,168,0,3,1,102]
s: 168
r: 158
i: [1,105,0,4,1,0,0,1] d3_subset_1
s: 105
r: 146
d: [2,196,0,3,1,1,23,146,105]
a: [3,200,0,3,1,196]
s: 200
r: 143
r: 153
d: [2,203,0,4,1,2,23,153,200]
a: [3,212,0,4,1,203]
s: 212
d: [2,202,0,5,1,2,23,143,200]
a: [3,207,0,5,1,202]
s: 207
i: [1,110,0,4,2,0,0,255] cc2_zfmisc_1
s: 110
i: [1,113,0,4,2,0,0,255] cc1_zfmisc_1
s: 113
d: [2,214,1,2,1,1,40,110,163]
a: [3,220,1,2,1,214]
s: 220
i: [1,121,0,4,1,0,0,1] existence_m1_subset_1
s: 121
d: [2,216,1,2,1,1,40,110,183]
a: [3,225,1,2,1,216]
s: 225
i: [1,106,0,5,1,0,0,255] t4_boole
s: 106
i: [1,107,0,5,1,0,0,255] t3_boole
s: 107
d: [2,215,1,3,1,0,40,110,123]
s: 215
i: [1,108,0,5,1,0,0,2] t1_boole
s: 108
i: [1,109,0,5,1,0,0,1] dt_k2_subset_1
s: 109
d: [2,227,0,4,1,0,22,109,105]
s: 227
r: 212
i: [1,112,0,5,2,0,0,1] rc1_subset_1
s: 112
i: [1,114,0,5,2,0,0,2] t6_boole
s: 114
i: [1,122,0,5,1,0,0,1] rc2_subset_1
s: 122
i: [1,125,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 125
i: [1,127,0,6,2,0,0,2] fc4_xboole_0
s: 127
i: [1,128,0,6,2,0,0,2] fc5_xboole_0
s: 128
d: [2,229,1,5,2,0,44,104,114]
s: 229
i: [1,111,0,7,2,0,0,1] rc1_subset_1
s: 111
d: [2,239,2,4,1,1,40,229,207]
a: [3,243,2,4,1,239]
s: 243
i: [1,120,0,7,3,0,0,2] t8_boole
s: 120
i: [1,126,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 126
d: [2,245,1,5,1,0,44,126,108]
s: 245
i: [1,115,0,8,2,0,0,1] t12_xboole_1
s: 115
d: [2,246,1,5,1,0,44,126,108]
s: 246
i: [1,118,0,8,2,0,0,255] t37_xboole_1
s: 118
i: [1,119,0,8,2,0,0,255] t37_xboole_1
s: 119
d: [2,247,1,5,1,0,44,108,126]
s: 247
i: [1,129,0,10,2,0,0,1] dt_k3_subset_1
s: 129
d: [2,250,1,5,1,0,44,108,126]
s: 250
i: [1,116,0,11,2,0,0,255] d4_subset_1
s: 116
i: [1,130,0,11,2,0,0,1] involutiveness_k3_subset_1
s: 130
d: [2,260,1,6,2,0,44,115,108]
s: 260
i: [1,117,0,13,2,0,0,1] t10_subset_1
s: 117
d: [2,263,1,6,2,0,44,108,115]
s: 263
d: [2,286,0,12,2,0,22,117,105]
s: 286
i: [1,131,0,14,3,0,0,2] idempotence_k4_subset_1
s: 131
d: [2,269,1,6,2,0,44,118,107]
s: 269
i: [1,132,0,15,3,0,0,2] dt_k4_subset_1
s: 132
d: [2,271,1,6,2,0,44,119,107]
s: 271
i: [1,133,0,16,3,0,0,2] redefinition_k4_subset_1
s: 133
i: [1,134,0,17,3,0,0,2] commutativity_k4_subset_1
s: 134
d: [2,272,1,6,2,0,44,107,119]
s: 272
d: [2,228,1,7,3,0,44,114,114]
s: 228
d: [2,230,1,7,2,0,44,106,114]
s: 230
d: [2,231,1,7,2,0,44,107,114]
s: 231
d: [2,232,1,7,2,0,44,108,114]
s: 232
d: [2,267,1,7,3,0,44,127,115]
s: 267
d: [2,261,1,8,2,0,44,115,126]
s: 261
d: [2,253,2,7,2,0,44,245,114]
s: 253
d: [2,262,1,8,2,0,44,115,126]
s: 262
d: [2,264,1,8,2,0,44,126,115]
s: 264
d: [2,330,2,7,2,0,44,232,126]
s: 330
d: [2,265,1,8,2,0,44,126,115]
s: 265
d: [2,331,2,7,2,0,44,232,126]
s: 331
d: [2,274,1,9,2,0,44,116,106]
s: 274
d: [2,275,1,9,2,0,44,116,107]
s: 275
d: [2,334,2,7,2,0,44,126,232]
s: 334
d: [2,277,1,9,2,0,44,106,116]
s: 277
d: [2,335,2,7,2,0,44,126,232]
s: 335
d: [2,278,1,9,2,0,44,107,116]
s: 278
d: [2,295,1,21,4,0,44,133,286]
d: [2,304,1,17,3,0,18,295]
d: [2,305,1,11,2,0,21,304,129]
s: 305
d: [2,365,2,13,3,0,44,274,130]
d: [2,373,2,7,2,0,21,365,129]
s: 373
d: [2,297,1,21,4,0,44,286,133]
d: [2,302,1,17,3,0,18,297]
d: [2,306,1,11,2,0,21,302,129]
s: 306
d: [2,407,3,4,1,0,40,373,122]
a: [3,417,3,4,1,407]
s: 417
d: [2,276,1,12,3,0,44,116,119]
s: 276
d: [2,425,3,9,3,1,44,243,276]
d: [2,440,3,7,2,2,21,425,163]
d: [2,441,3,3,1,2,21,440,227]
d: [2,442,3,0,0,2,21,441,124]
e: 442
d: [2,426,2,10,3,1,44,207,276]
d: [2,444,2,7,2,1,21,426,104]
d: [2,445,2,3,1,1,21,444,227]
d: [2,446,2,0,0,1,21,445,124]
e: 446
r: 200
r: 207
r: 243
a: [3,199,0,3,1,196]
s: 199
s: 153
a: [3,142,0,5,1,99]
s: 142
d: [2,452,2,6,2,1,40,142,267]
a: [3,457,2,2,1,452]
s: 457
a: [3,206,0,5,1,202]
s: 206
d: [2,279,1,12,3,0,44,119,116]
s: 279
a: [3,242,2,4,1,239]
s: 242
d: [2,280,1,12,3,0,44,118,116]
s: 280
d: [2,470,3,4,1,1,40,242,110]
a: [3,475,3,4,1,470]
s: 475
d: [2,287,1,18,3,0,44,286,130]
d: [2,288,1,12,2,0,21,287,129]
s: 288
d: [2,299,1,22,5,0,44,132,133]
d: [2,300,1,14,3,0,18,299]
s: 300
d: [2,406,3,5,1,0,40,373,121]
a: [3,412,3,5,1,406]
s: 412
d: [2,311,1,24,5,0,44,133,134]
d: [2,314,1,16,3,0,18,311]
s: 314
d: [2,419,4,6,2,1,44,417,114]
s: 419
d: [2,309,1,24,5,0,44,133,134]
d: [2,316,1,16,3,0,18,309]
s: 316
d: [2,308,1,24,5,0,44,134,133]
d: [2,317,1,16,3,0,18,308]
s: 317
d: [2,369,2,13,3,0,44,130,274]
d: [2,374,2,7,2,0,21,369,129]
s: 374
d: [2,307,1,24,5,0,44,134,133]
d: [2,318,1,16,3,0,18,307]
s: 318
d: [2,384,2,13,3,0,44,305,274]
d: [2,397,2,9,2,0,18,384]
d: [2,401,2,7,2,0,22,397,108]
s: 401
d: [2,285,2,8,3,0,44,260,114]
s: 285
d: [2,463,2,6,2,1,40,206,267]
a: [3,468,2,2,1,463]
s: 468
d: [2,524,4,7,2,1,44,412,114]
s: 524
d: [2,289,2,8,3,0,44,269,114]
s: 289
d: [2,325,2,8,3,0,44,231,119]
s: 325
d: [2,327,2,8,3,0,44,119,231]
s: 327
d: [2,328,2,8,3,0,44,118,231]
s: 328
d: [2,329,2,8,3,0,44,232,115]
s: 329
d: [2,333,2,8,3,0,44,115,232]
s: 333
d: [2,391,2,8,3,0,44,127,305]
s: 391
d: [2,342,2,9,3,0,44,261,115]
s: 342
d: [2,346,2,9,3,0,44,115,261]
s: 346
d: [2,589,3,10,2,1,40,342,206]
a: [3,594,3,5,1,589]
s: 594
d: [2,367,2,13,3,0,44,130,274]
d: [2,371,2,9,2,0,18,367]
s: 371
d: [2,606,4,10,3,1,44,594,276]
d: [2,607,4,6,2,1,21,606,227]
d: [2,608,4,3,1,1,21,607,124]
a: [3,612,4,3,1,608]
s: 612
d: [2,377,2,13,3,0,44,130,275]
d: [2,380,2,9,2,0,18,377]
s: 380
d: [2,366,2,14,3,0,44,286,274]
d: [2,372,2,10,2,0,18,366]
s: 372
d: [2,590,3,10,2,1,40,342,142]
a: [3,599,3,5,1,590]
s: 599
d: [2,376,2,14,3,0,44,286,275]
d: [2,381,2,10,2,0,18,376]
s: 381
d: [2,681,4,10,3,1,44,599,276]
d: [2,682,4,6,2,2,21,681,153]
a: [3,686,4,3,1,682]
s: 686
d: [2,392,2,10,3,0,44,128,305]
s: 392
d: [2,427,2,16,4,0,44,305,276]
d: [2,437,2,12,3,0,18,427]
d: [2,448,2,10,3,0,22,437,108]
s: 448
d: [2,633,5,5,2,1,44,612,114]
s: 633
d: [2,490,2,14,3,0,44,288,274]
d: [2,505,2,10,2,0,18,490]
s: 505
d: [2,584,3,6,2,0,40,391,121]
s: 584
d: [2,488,2,14,3,0,44,288,275]
d: [2,507,2,10,2,0,18,488]
s: 507
d: [2,319,2,11,3,0,44,230,116]
s: 319
d: [2,716,3,6,2,0,40,392,227]
s: 716
d: [2,321,2,11,3,0,44,116,230]
s: 321
d: [2,761,4,6,2,0,40,584,110]
s: 761
d: [2,324,2,11,3,0,44,231,116]
s: 324
d: [2,326,2,11,3,0,44,116,231]
s: 326
d: [2,814,4,6,2,0,40,716,110]
s: 814
d: [2,364,2,11,3,0,44,274,114]
s: 364
d: [2,862,5,8,3,0,44,814,371]
d: [2,866,5,6,2,1,21,862,163]
s: 866
d: [2,375,2,11,3,0,44,275,114]
s: 375
d: [2,387,2,11,3,0,44,305,232]
s: 387
d: [2,723,3,7,2,0,40,392,122]
s: 723
d: [2,393,2,11,3,0,44,232,305]
s: 393
d: [2,892,4,7,2,0,40,723,110]
s: 892
d: [2,430,2,15,4,0,44,129,276]
d: [2,434,2,11,3,0,18,430]
s: 434
d: [2,386,2,12,3,0,44,305,261]
s: 386
d: [2,918,3,0,0,3,52,199,153,142,386]
e: 918
d: [2,921,3,7,2,1,40,386,142]
d: [2,936,3,3,1,2,21,921,153]
d: [2,937,3,0,0,3,21,936,199]
e: 937
i: [1,98,0,9,2,1,0,0] t19_subset_1
a: [3,147,0,4,1,99]
d: [2,450,0,3,1,1,22,147,105]
f: 937,450,446,202,196,97,98,99
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t19_subset_1
% SZS output start Proof for subset_1__t19_subset_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => (r1_tarski(k3_subset_1(X0,X1),X1) <=> k2_subset_1(X0) = X1)) [input t19_subset_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => (r1_tarski(k3_subset_1(X0,X1),X1) <=> k2_subset_1(X0) = X1)) [negated conjecture 1]
6. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
9. ! [X0] : k2_subset_1(X0) = X0 [input d3_subset_1]
10. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => k3_subset_1(X0,X1) = k4_xboole_0(X0,X1)) [input d4_subset_1]
13. ! [X0] : m1_subset_1(k2_subset_1(X0),k1_zfmisc_1(X0)) [input dt_k2_subset_1]
15. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => m1_subset_1(k3_subset_1(X0,X1),k1_zfmisc_1(X0))) [input dt_k3_subset_1]
34. ! [X0,X1,X2] : ((m1_subset_1(X2,k1_zfmisc_1(X0)) & m1_subset_1(X1,k1_zfmisc_1(X0))) => k4_subset_1(X0,X1,X2) = k2_xboole_0(X1,X2)) [input redefinition_k4_subset_1]
35. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
36. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => k2_subset_1(X0) = k4_subset_1(X0,X1,k3_subset_1(X0,X1))) [input t10_subset_1]
37. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k2_xboole_0(X0,X1) = X1) [input t12_xboole_1]
39. ! [X0] : r1_tarski(k1_xboole_0,X0) [input t2_xboole_1]
40. ! [X0] : ! [X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [input t37_xboole_1]
46. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => (r1_tarski(k3_subset_1(X0,X1),X1) <=> k2_subset_1(X0) = X1)) [flattening 2]
47. ! [X0,X1] : (r1_tarski(X0,X1) => k2_xboole_0(X0,X1) = X1) [flattening 37]
48. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => k3_subset_1(X0,X1) = k4_xboole_0(X0,X1)) [flattening 10]
49. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => k2_subset_1(X0) = k4_subset_1(X0,X1,k3_subset_1(X0,X1))) [flattening 36]
50. ! [X0,X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [flattening 40]
53. ! [X0] : r1_tarski(X0,X0) [rectify 35]
57. ? [X0,X1] : ((r1_tarski(k3_subset_1(X0,X1),X1) <~> k2_subset_1(X0) = X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 46]
62. ! [X0,X1] : (k2_xboole_0(X0,X1) = X1 | ~r1_tarski(X0,X1)) [ennf transformation 47]
63. ! [X0,X1] : (k3_subset_1(X0,X1) = k4_xboole_0(X0,X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 48]
64. ! [X0,X1] : (k2_subset_1(X0) = k4_subset_1(X0,X1,k3_subset_1(X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 49]
68. ! [X0,X1] : (m1_subset_1(k3_subset_1(X0,X1),k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 15]
74. ! [X0,X1,X2] : (k4_subset_1(X0,X1,X2) = k2_xboole_0(X1,X2) | (~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)))) [ennf transformation 34]
75. ! [X0,X1,X2] : (k4_subset_1(X0,X1,X2) = k2_xboole_0(X1,X2) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 74]
78. ? [X0,X1] : (((k2_subset_1(X0) != X1 | ~r1_tarski(k3_subset_1(X0,X1),X1)) & (k2_subset_1(X0) = X1 | r1_tarski(k3_subset_1(X0,X1),X1))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [nnf transformation 57]
79. ? [X0,X1] : ((k2_subset_1(X0) != X1 | ~r1_tarski(k3_subset_1(X0,X1),X1)) & (k2_subset_1(X0) = X1 | r1_tarski(k3_subset_1(X0,X1),X1)) & m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 78]
80. ? [X0,X1] : ((k2_subset_1(X0) != X1 | ~r1_tarski(k3_subset_1(X0,X1),X1)) & (k2_subset_1(X0) = X1 | r1_tarski(k3_subset_1(X0,X1),X1)) & m1_subset_1(X1,k1_zfmisc_1(X0))) => ((sK1 != k2_subset_1(sK0) | ~r1_tarski(k3_subset_1(sK0,sK1),sK1)) & (sK1 = k2_subset_1(sK0) | r1_tarski(k3_subset_1(sK0,sK1),sK1)) & m1_subset_1(sK1,k1_zfmisc_1(sK0))) [choice axiom]
81. (sK1 != k2_subset_1(sK0) | ~r1_tarski(k3_subset_1(sK0,sK1),sK1)) & (sK1 = k2_subset_1(sK0) | r1_tarski(k3_subset_1(sK0,sK1),sK1)) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) [skolemisation 79,80]
84. ! [X0,X1] : ((k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | k1_xboole_0 != k4_xboole_0(X0,X1))) [nnf transformation 50]
97. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 81] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
98. sK1 = k2_subset_1(sK0) | r1_tarski(k3_subset_1(sK0,sK1),sK1) [cnf transformation 81] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
99. sK1 != k2_subset_1(sK0) | ~r1_tarski(k3_subset_1(sK0,sK1),sK1) [cnf transformation 81] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
104. r1_tarski(k1_xboole_0,X0) [cnf transformation 39] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
105. k2_subset_1(X0) = X0 [cnf transformation 9] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
109. m1_subset_1(k2_subset_1(X0),k1_zfmisc_1(X0)) [cnf transformation 13] {a:0,w:5,wCS:2000,sine:1,thAx:0,allAx:1,thDist:-1}
115. k2_xboole_0(X0,X1) = X1 | ~r1_tarski(X0,X1) [cnf transformation 62] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
116. k3_subset_1(X0,X1) = k4_xboole_0(X0,X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 63] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
117. k2_subset_1(X0) = k4_subset_1(X0,X1,k3_subset_1(X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 64] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:1,thDist:-1}
119. k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1) [cnf transformation 84] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
124. r1_tarski(X0,X0) [cnf transformation 53] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
126. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 6] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
129. m1_subset_1(k3_subset_1(X0,X1),k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 68] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
133. k4_subset_1(X0,X1,X2) = k2_xboole_0(X1,X2) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 75] {a:0,w:16,wCS:6400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
141. 1 <=> r1_tarski(k3_subset_1(sK0,sK1),sK1) [avatar definition]
142. r1_tarski(k3_subset_1(sK0,sK1),sK1) <- (1) [avatar component clause 141] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
143. ~r1_tarski(k3_subset_1(sK0,sK1),sK1) <- (~1) [avatar component clause 141] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
145. 2 <=> sK1 = k2_subset_1(sK0) [avatar definition]
146. sK1 = k2_subset_1(sK0) <- (2) [avatar component clause 145] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
147. sK1 != k2_subset_1(sK0) <- (~2) [avatar component clause 145] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
148. ~1 | ~2 [avatar split clause 99,145,141]
149. 1 | 2 [avatar split clause 98,145,141]
151. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
153. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (3) [avatar component clause 151] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
154. 3 [avatar split clause 97,151]
196. sK0 = sK1 <- (2) [backward demodulation 146,105] {a:0,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
198. 12 <=> sK0 = sK1 [avatar definition]
199. sK0 != sK1 <- (~12) [avatar component clause 198] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
200. sK0 = sK1 <- (12) [avatar component clause 198] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
201. 12 | ~2 [avatar split clause 196,145,198]
202. ~r1_tarski(k3_subset_1(sK0,sK0),sK0) <- (~1, 12) [backward demodulation 143,200] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
205. 13 <=> r1_tarski(k3_subset_1(sK0,sK0),sK0) [avatar definition]
207. ~r1_tarski(k3_subset_1(sK0,sK0),sK0) <- (~13) [avatar component clause 205] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
208. ~13 | 1 | ~12 [avatar split clause 202,198,141,205]
227. m1_subset_1(X0,k1_zfmisc_1(X0)) [forward demodulation 109,105] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
261. k2_xboole_0(X5,X4) = X5 | ~r1_tarski(X4,X5) [superposition 115,126] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
276. k1_xboole_0 = k3_subset_1(X4,X5) | ~m1_subset_1(X5,k1_zfmisc_1(X4)) | ~r1_tarski(X4,X5) [superposition 116,119] {a:1,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
286. k4_subset_1(X0,X1,k3_subset_1(X0,X1)) = X0 | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [forward demodulation 117,105] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
295. k2_xboole_0(X1,k3_subset_1(X0,X1)) = X0 | ~m1_subset_1(k3_subset_1(X0,X1),k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [superposition 133,286] {a:1,w:21,wCS:8400,sine:1,thAx:0,allAx:2,thDist:-2}
304. k2_xboole_0(X1,k3_subset_1(X0,X1)) = X0 | ~m1_subset_1(k3_subset_1(X0,X1),k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [duplicate literal removal 295] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
305. k2_xboole_0(X1,k3_subset_1(X0,X1)) = X0 | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [subsumption resolution 304,129] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
386. ~r1_tarski(k3_subset_1(X5,X4),X4) | ~m1_subset_1(X4,k1_zfmisc_1(X5)) | X4 = X5 [superposition 305,261] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
426. ~r1_tarski(k1_xboole_0,sK0) | ~m1_subset_1(sK0,k1_zfmisc_1(sK0)) | ~r1_tarski(sK0,sK0) <- (~13) [superposition 207,276] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
444. ~m1_subset_1(sK0,k1_zfmisc_1(sK0)) | ~r1_tarski(sK0,sK0) <- (~13) [subsumption resolution 426,104] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
445. ~r1_tarski(sK0,sK0) <- (~13) [subsumption resolution 444,227] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
446. $false <- (~13) [subsumption resolution 445,124] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
447. 13 [avatar contradiction clause 446]
450. sK0 != sK1 <- (~2) [forward demodulation 147,105] {a:0,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
451. ~12 | 2 [avatar split clause 450,145,198]
921. ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) | sK0 = sK1 <- (1) [resolution 386,142] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
936. sK0 = sK1 <- (1, 3) [subsumption resolution 921,153] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
937. $false <- (1, 3, ~12) [subsumption resolution 936,199] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
938. ~1 | ~3 | 12 [avatar contradiction clause 937]
940. $false [avatar sat refutation 148,149,154,201,208,447,451,938] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:168,thDist:-168}
% SZS output end Proof for subset_1__t19_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 44
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 43
% Generated clauses: 724
% Active clauses: 111
% Passive clauses: 218
% Final active clauses: 103
% Final passive clauses: 72
% Main loop iterations started: 146
% 
% >>> Simplifying Inferences
% Duplicate literals: 146
% Trivial inequalities: 3
% Fw subsumption resolutions: 52
% Fw demodulations: 20
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 45
% Equational tautologies: 10
% Forward subsumptions: 247
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 89
% Unit resulting resolution: 1
% Forward superposition: 170
% Backward superposition: 174
% Self superposition: 2
% Equality resolution: 3
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 16
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 34
% SAT solver unit clauses: 12
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.002 s
% forward subsumption index maintenance: 0.001 s ( own 0.0000 s ) 
% resolution: 0.001 s
% superposition: 0.002 s ( own 0.001 s ) 
% other: 0.004 s ( own 0.0000 s ) 

