i: [1,29,0,2,1,1,0,0] t17_xxreal_0
a: [3,47,0,2,1,29]
s: 47
i: [1,28,0,2,1,1,0,0] t17_xxreal_0
a: [3,52,0,2,1,28]
s: 52
i: [1,30,0,5,1,1,0,0] t17_xxreal_0
a: [3,42,0,5,1,30]
s: 42
i: [1,38,0,2,1,0,0,1] rc1_xxreal_0
a: [3,68,0,2,1,38]
s: 68
i: [1,33,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,59,0,5,2,33]
s: 59
i: [1,35,0,9,3,0,0,255] idempotence_k3_xxreal_0
a: [3,63,0,7,2,35]
s: 63
i: [1,34,0,8,3,0,0,1] fc5_xxreal_0
s: 34
i: [1,37,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 37
d: [2,72,1,11,3,1,40,37,42]
d: [2,75,1,9,2,2,21,72,52]
a: [3,79,1,4,1,75]
s: 79
d: [2,85,2,0,0,3,52,52,47,79,34]
e: 85
d: [2,86,2,4,2,1,40,79,34]
d: [2,88,2,2,1,2,21,86,47]
d: [2,89,2,0,0,3,21,88,52]
e: 89
r: 79
i: [1,36,0,11,3,0,0,255] commutativity_k3_xxreal_0
s: 36
i: [1,31,0,12,4,0,0,255] d9_xxreal_0
s: 31
d: [2,97,1,10,4,1,44,42,31]
d: [2,106,1,8,3,2,21,97,47]
d: [2,107,1,6,2,3,21,106,52]
a: [3,111,1,3,1,107]
s: 111
r: 111
i: [1,32,0,12,4,0,0,255] d9_xxreal_0
s: 32
a: [3,110,1,3,1,107]
s: 110
a: [3,115,1,3,1,107]
s: 115
d: [2,143,2,0,0,3,52,52,115,59]
e: 143
d: [2,144,2,0,0,2,52,52,52,115,115,37]
e: 144
d: [2,145,2,0,0,2,52,52,52,115,115,37]
e: 145
d: [2,146,2,2,1,2,40,115,59]
d: [2,152,2,0,0,3,21,146,52]
e: 152
d: [2,117,2,7,3,1,40,111,37]
d: [2,118,2,5,2,2,21,117,52]
d: [2,119,2,3,1,3,21,118,47]
d: [2,128,1,10,4,1,44,42,32]
d: [2,140,1,8,3,2,21,128,47]
d: [2,141,1,6,2,3,21,140,52]
f: 145,141,119,107,28,29,30
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_0__t17_xxreal_0
% SZS output start Proof for xxreal_0__t17_xxreal_0
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => r1_xxreal_0(k3_xxreal_0(X0,X1),X0))) [input t17_xxreal_0]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => r1_xxreal_0(k3_xxreal_0(X0,X1),X0))) [negated conjecture 1]
4. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
5. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((~r1_xxreal_0(X0,X1) => k3_xxreal_0(X0,X1) = X1) & (r1_xxreal_0(X0,X1) => k3_xxreal_0(X0,X1) = X0)))) [input d9_xxreal_0]
11. ? [X0] : (? [X1] : (~r1_xxreal_0(k3_xxreal_0(X0,X1),X0) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
12. ! [X0] : (! [X1] : (((k3_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X0,X1)) & (k3_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 5]
21. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 4]
22. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 21]
23. ? [X0] : (? [X1] : (~r1_xxreal_0(k3_xxreal_0(X0,X1),X0) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (~r1_xxreal_0(k3_xxreal_0(sK0,X1),sK0) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
24. ? [X1] : (~r1_xxreal_0(k3_xxreal_0(sK0,X1),sK0) & v1_xxreal_0(X1)) => (~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) & v1_xxreal_0(sK1)) [choice axiom]
25. (~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 11,24,23]
28. v1_xxreal_0(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
29. v1_xxreal_0(sK1) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
30. ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) [cnf transformation 25] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
31. k3_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 12] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
32. k3_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 12] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
37. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 22] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
40. 1 <=> r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) [avatar definition]
42. ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) <- (~1) [avatar component clause 40] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
43. ~1 [avatar split clause 30,40]
45. 2 <=> v1_xxreal_0(sK1) [avatar definition]
47. v1_xxreal_0(sK1) <- (2) [avatar component clause 45] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
48. 2 [avatar split clause 29,45]
50. 3 <=> v1_xxreal_0(sK0) [avatar definition]
52. v1_xxreal_0(sK0) <- (3) [avatar component clause 50] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
53. 3 [avatar split clause 28,50]
97. ~r1_xxreal_0(sK0,sK0) | ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~1) [superposition 42,31] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
106. ~r1_xxreal_0(sK0,sK0) | ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK0) <- (~1, 2) [subsumption resolution 97,47] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
107. ~r1_xxreal_0(sK0,sK0) | ~r1_xxreal_0(sK0,sK1) <- (~1, 2, 3) [subsumption resolution 106,52] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
109. 10 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
111. ~r1_xxreal_0(sK0,sK1) <- (~10) [avatar component clause 109] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
113. 11 <=> r1_xxreal_0(sK0,sK0) [avatar definition]
115. ~r1_xxreal_0(sK0,sK0) <- (~11) [avatar component clause 113] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
116. ~10 | ~11 | 1 | ~2 | ~3 [avatar split clause 107,50,45,40,113,109]
117. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~10) [resolution 111,37] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
118. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) <- (3, ~10) [subsumption resolution 117,52] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
119. r1_xxreal_0(sK1,sK0) <- (2, 3, ~10) [subsumption resolution 118,47] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
121. 12 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
124. 12 | ~2 | ~3 | 10 [avatar split clause 119,109,50,45,121]
128. ~r1_xxreal_0(sK1,sK0) | r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~1) [superposition 42,32] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
140. ~r1_xxreal_0(sK1,sK0) | r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK0) <- (~1, 2) [subsumption resolution 128,47] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
141. ~r1_xxreal_0(sK1,sK0) | r1_xxreal_0(sK0,sK1) <- (~1, 2, 3) [subsumption resolution 140,52] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
142. 10 | ~12 | 1 | ~2 | ~3 [avatar split clause 141,50,45,40,121,109]
145. $false <- (3, ~11) [unit resulting resolution 52,52,115,115,37] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
151. ~3 | 11 [avatar contradiction clause 145]
154. $false [avatar sat refutation 43,48,53,116,124,142,151] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:115,thDist:-115}
% SZS output end Proof for xxreal_0__t17_xxreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 11
% Generated clauses: 77
% Active clauses: 15
% Passive clauses: 25
% Final active clauses: 13
% Final passive clauses: 10
% Main loop iterations started: 15
% 
% >>> Simplifying Inferences
% Duplicate literals: 32
% Fw subsumption resolutions: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 5
% Unit resulting resolution: 4
% Factoring: 1
% Forward superposition: 6
% Backward superposition: 11
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 4
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

