i: [1,112,0,2,1,1,0,0] t31_diraf
a: [3,182,0,2,1,112]
s: 182
i: [1,111,0,2,1,1,0,0] t31_diraf
a: [3,187,0,2,1,111]
s: 187
i: [1,110,0,2,1,1,0,0] t31_diraf
a: [3,192,0,2,1,110]
s: 192
i: [1,114,0,4,1,1,0,0] t31_diraf
a: [3,172,0,4,1,114]
s: 172
i: [1,113,0,4,1,1,0,0] t31_diraf
a: [3,177,0,4,1,113]
s: 177
i: [1,115,0,15,3,1,0,0] t31_diraf
a: [3,159,0,5,1,115]
s: 159
i: [1,116,0,2,1,0,0,255] spc1_boole
a: [3,197,0,2,1,116]
s: 197
i: [1,117,0,2,1,0,0,255] fc1_xboole_0
a: [3,202,0,2,1,117]
s: 202
i: [1,152,0,2,1,0,0,2] rc2_xboole_0
a: [3,207,0,2,1,152]
s: 207
i: [1,153,0,2,1,0,0,1] existence_l1_analoaf
a: [3,212,0,2,1,153]
s: 212
i: [1,154,0,2,1,0,0,2] existence_l1_struct_0
a: [3,217,0,2,1,154]
s: 217
i: [1,155,0,2,1,0,0,2] rc1_xboole_0
a: [3,222,0,2,1,155]
s: 222
i: [1,118,0,4,2,0,0,1] dt_l1_analoaf
s: 118
d: [2,224,1,2,1,1,40,118,182]
a: [3,229,1,2,1,224]
s: 229
i: [1,150,0,4,1,0,0,1] existence_m1_subset_1
s: 150
i: [1,133,0,5,2,0,0,255] t6_boole
s: 133
d: [2,225,1,2,1,1,40,118,212]
a: [3,234,1,2,1,225]
s: 234
i: [1,149,0,5,2,0,0,255] t7_boole
s: 149
i: [1,119,0,6,3,0,0,255] cc7_struct_0
s: 119
i: [1,120,0,6,3,0,0,255] cc5_struct_0
s: 120
i: [1,122,0,6,3,0,0,1] cc2_struct_0
s: 122
d: [2,238,1,4,2,1,40,119,217]
a: [3,244,1,2,1,238]
s: 244
i: [1,123,0,6,3,0,0,1] cc17_struct_0
s: 123
d: [2,240,2,4,2,1,40,119,234]
a: [3,253,2,2,1,240]
s: 253
i: [1,124,0,6,3,0,0,1] cc1_struct_0
s: 124
i: [1,127,0,6,3,0,0,255] cc4_struct_0
s: 127
i: [1,128,0,6,3,0,0,255] cc6_struct_0
s: 128
i: [1,146,0,6,2,0,0,255] t1_subset
s: 146
i: [1,151,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 151
i: [1,125,0,7,3,0,0,255] cc8_struct_0
s: 125
i: [1,129,0,7,3,0,0,255] cc11_struct_0
s: 129
i: [1,130,0,7,3,0,0,255] cc11_struct_0
s: 130
i: [1,131,0,7,3,0,0,255] cc9_struct_0
s: 131
i: [1,134,0,7,3,0,0,255] fc9_struct_0
s: 134
i: [1,137,0,7,3,0,0,1] fc2_struct_0
s: 137
i: [1,142,0,7,3,0,0,255] fc6_struct_0
s: 142
i: [1,143,0,7,3,0,0,255] fc8_struct_0
s: 143
i: [1,144,0,7,3,0,0,1] fc1_struct_0
s: 144
i: [1,145,0,7,3,0,0,255] fc7_struct_0
s: 145
i: [1,148,0,7,3,0,0,255] t8_boole
s: 148
i: [1,147,0,8,3,0,0,255] t2_subset
s: 147
d: [2,269,1,7,2,1,40,147,172]
a: [3,279,1,3,1,269]
s: 279
r: 279
i: [1,132,0,9,4,0,0,255] cc10_struct_0
s: 132
a: [3,278,1,3,1,269]
s: 278
i: [1,138,0,20,6,0,0,1] t19_diraf
s: 138
i: [1,139,0,20,6,0,0,1] t19_diraf
s: 139
d: [2,286,2,4,2,1,40,279,137]
d: [2,287,2,2,1,2,21,286,229]
a: [3,290,2,2,1,287]
s: 290
i: [1,140,0,24,7,0,0,1] t20_diraf
s: 140
a: [3,275,1,4,1,269]
s: 275
i: [1,141,0,24,7,0,0,1] t20_diraf
s: 141
i: [1,135,0,27,7,0,0,1] d5_diraf
s: 135
d: [2,270,1,7,2,1,40,147,177]
a: [3,284,1,4,1,270]
s: 284
i: [1,136,0,27,7,0,0,1] d5_diraf
s: 136
d: [2,300,2,4,1,1,40,275,151]
a: [3,306,2,4,1,300]
s: 306
d: [2,271,1,6,2,0,40,147,150]
s: 271
d: [2,236,1,7,3,0,44,133,133]
s: 236
d: [2,308,2,4,1,1,40,284,151]
a: [3,314,2,4,1,308]
s: 314
d: [2,259,1,9,4,0,44,125,133]
s: 259
d: [2,326,2,6,2,0,40,271,151]
s: 326
d: [2,261,1,9,4,0,44,131,133]
s: 261
d: [2,316,1,35,11,0,40,136,138]
d: [2,325,1,17,6,0,18,316]
s: 325
d: [2,320,1,39,12,0,40,136,141]
d: [2,321,1,21,7,0,18,320]
s: 321
d: [2,341,3,0,0,7,52,192,187,290,182,177,172,159,321]
e: 341
d: [2,342,2,16,6,1,40,321,159]
d: [2,344,2,12,5,2,21,342,172]
d: [2,345,2,8,4,3,21,344,177]
d: [2,346,2,6,3,4,21,345,182]
d: [2,347,2,4,2,5,21,346,290]
d: [2,348,2,2,1,6,21,347,187]
d: [2,349,2,0,0,7,21,348,192]
e: 349
r: 278
r: 290
r: 275
r: 284
s: 279
a: [3,291,2,2,1,287]
s: 291
d: [2,359,3,0,0,3,52,229,192,291,122]
e: 359
d: [2,360,3,4,2,1,40,291,122]
d: [2,363,3,2,1,2,21,360,192]
d: [2,364,3,0,0,3,21,363,229]
e: 364
r: 159
r: 279
r: 291
a: [3,158,0,5,1,115]
s: 158
a: [3,163,0,5,1,115]
s: 163
s: 290
s: 278
s: 275
s: 284
d: [2,319,1,39,12,0,40,136,140]
d: [2,322,1,21,7,0,18,319]
s: 322
d: [2,317,1,35,11,0,40,136,139]
d: [2,324,1,21,7,0,18,317]
s: 324
d: [2,374,3,0,0,7,52,192,187,290,182,177,172,163,324]
e: 374
d: [2,375,2,16,6,1,40,324,163]
d: [2,377,2,12,5,2,21,375,172]
d: [2,378,2,8,4,3,21,377,177]
d: [2,379,2,6,3,4,21,378,182]
d: [2,380,2,4,2,5,21,379,290]
d: [2,381,2,2,1,6,21,380,187]
d: [2,382,2,0,0,7,21,381,192]
e: 382
r: 163
a: [3,162,0,5,1,115]
s: 162
a: [3,167,0,5,1,115]
s: 167
d: [2,384,3,0,0,7,52,192,187,290,182,172,177,167,322]
e: 384
d: [2,385,2,16,6,1,40,167,322]
d: [2,387,2,12,5,2,21,385,177]
d: [2,388,2,8,4,3,21,387,172]
d: [2,389,2,6,3,4,21,388,182]
d: [2,390,2,4,2,5,21,389,290]
d: [2,391,2,2,1,6,21,390,187]
d: [2,392,2,0,0,7,21,391,192]
e: 392
d: [2,353,2,4,2,5,21,346,187]
d: [2,354,2,2,1,6,21,353,192]
f: 392,382,364,354,224,110,111,112,113,114,115
% Refutation found. Thanks to Tanya!
% SZS status Theorem for diraf__t31_diraf
% SZS output start Proof for diraf__t31_diraf
1. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r3_diraf(X0,X1,X2,X1) & r3_diraf(X0,X1,X2,X2) & r3_diraf(X0,X1,X1,X2))))) [input t31_diraf]
2. ~! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r3_diraf(X0,X1,X2,X1) & r3_diraf(X0,X1,X2,X2) & r3_diraf(X0,X1,X1,X2))))) [negated conjecture 1]
8. ! [X0] : (l1_struct_0(X0) => (~v7_struct_0(X0) => ~v2_struct_0(X0))) [input cc2_struct_0]
15. ! [X0] : ((l1_analoaf(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r3_diraf(X0,X1,X2,X3) <=> r2_diraf(X0,X1,X2,X1,X3)))))) [input d5_diraf]
17. ! [X0] : (l1_analoaf(X0) => l1_struct_0(X0)) [input dt_l1_analoaf]
34. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X1,X2) & r2_diraf(X0,X1,X2,X2,X1))))) [input t19_diraf]
36. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r2_diraf(X0,X3,X3,X1,X2) & r2_diraf(X0,X1,X2,X3,X3)))))) [input t20_diraf]
45. ? [X0] : (? [X1] : (? [X2] : ((~r3_diraf(X0,X1,X2,X1) | ~r3_diraf(X0,X1,X2,X2) | ~r3_diraf(X0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
46. ? [X0] : (? [X1] : (? [X2] : ((~r3_diraf(X0,X1,X2,X1) | ~r3_diraf(X0,X1,X2,X2) | ~r3_diraf(X0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) [flattening 45]
47. ! [X0] : (l1_struct_0(X0) | ~l1_analoaf(X0)) [ennf transformation 17]
52. ! [X0] : ((~v2_struct_0(X0) | v7_struct_0(X0)) | ~l1_struct_0(X0)) [ennf transformation 8]
53. ! [X0] : (~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0)) [flattening 52]
73. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r3_diraf(X0,X1,X2,X3) <=> r2_diraf(X0,X1,X2,X1,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | v2_struct_0(X0))) [ennf transformation 15]
74. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r3_diraf(X0,X1,X2,X3) <=> r2_diraf(X0,X1,X2,X1,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 73]
77. ! [X0] : (! [X1] : (! [X2] : ((r2_diraf(X0,X1,X2,X1,X2) & r2_diraf(X0,X1,X2,X2,X1)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0))) [ennf transformation 34]
78. ! [X0] : (! [X1] : (! [X2] : ((r2_diraf(X0,X1,X2,X1,X2) & r2_diraf(X0,X1,X2,X2,X1)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0)) [flattening 77]
79. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_diraf(X0,X3,X3,X1,X2) & r2_diraf(X0,X1,X2,X3,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0))) [ennf transformation 36]
80. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_diraf(X0,X3,X3,X1,X2) & r2_diraf(X0,X1,X2,X3,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0)) [flattening 79]
95. ? [X0] : (? [X1] : (? [X2] : ((~r3_diraf(X0,X1,X2,X1) | ~r3_diraf(X0,X1,X2,X2) | ~r3_diraf(X0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : ((~r3_diraf(sK0,X1,X2,X1) | ~r3_diraf(sK0,X1,X2,X2) | ~r3_diraf(sK0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
96. ? [X1] : (? [X2] : ((~r3_diraf(sK0,X1,X2,X1) | ~r3_diraf(sK0,X1,X2,X2) | ~r3_diraf(sK0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : ((~r3_diraf(sK0,sK1,X2,sK1) | ~r3_diraf(sK0,sK1,X2,X2) | ~r3_diraf(sK0,sK1,sK1,X2)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
97. ? [X2] : ((~r3_diraf(sK0,sK1,X2,sK1) | ~r3_diraf(sK0,sK1,X2,X2) | ~r3_diraf(sK0,sK1,sK1,X2)) & m1_subset_1(X2,u1_struct_0(sK0))) => ((~r3_diraf(sK0,sK1,sK2,sK1) | ~r3_diraf(sK0,sK1,sK2,sK2) | ~r3_diraf(sK0,sK1,sK1,sK2)) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
98. (((~r3_diraf(sK0,sK1,sK2,sK1) | ~r3_diraf(sK0,sK1,sK2,sK2) | ~r3_diraf(sK0,sK1,sK1,sK2)) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0) [skolemisation 46,97,96,95]
99. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r3_diraf(X0,X1,X2,X3) | ~r2_diraf(X0,X1,X2,X1,X3)) & (r2_diraf(X0,X1,X2,X1,X3) | ~r3_diraf(X0,X1,X2,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [nnf transformation 74]
110. ~v7_struct_0(sK0) [cnf transformation 98] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
111. v2_analoaf(sK0) [cnf transformation 98] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
112. l1_analoaf(sK0) [cnf transformation 98] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
113. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 98] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
114. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 98] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
115. ~r3_diraf(sK0,sK1,sK2,sK1) | ~r3_diraf(sK0,sK1,sK2,sK2) | ~r3_diraf(sK0,sK1,sK1,sK2) [cnf transformation 98] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
118. ~l1_analoaf(X0) | l1_struct_0(X0) [cnf transformation 47] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
122. ~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 53] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
136. ~r2_diraf(X0,X1,X2,X1,X3) | r3_diraf(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 99] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
139. r2_diraf(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 78] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
140. r2_diraf(X0,X1,X2,X3,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 80] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
141. r2_diraf(X0,X3,X3,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 80] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
157. 1 <=> r3_diraf(sK0,sK1,sK1,sK2) [avatar definition]
159. ~r3_diraf(sK0,sK1,sK1,sK2) <- (~1) [avatar component clause 157] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
161. 2 <=> r3_diraf(sK0,sK1,sK2,sK2) [avatar definition]
163. ~r3_diraf(sK0,sK1,sK2,sK2) <- (~2) [avatar component clause 161] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
165. 3 <=> r3_diraf(sK0,sK1,sK2,sK1) [avatar definition]
167. ~r3_diraf(sK0,sK1,sK2,sK1) <- (~3) [avatar component clause 165] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
168. ~1 | ~2 | ~3 [avatar split clause 115,165,161,157]
170. 4 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
172. m1_subset_1(sK2,u1_struct_0(sK0)) <- (4) [avatar component clause 170] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
173. 4 [avatar split clause 114,170]
175. 5 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
177. m1_subset_1(sK1,u1_struct_0(sK0)) <- (5) [avatar component clause 175] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
178. 5 [avatar split clause 113,175]
180. 6 <=> l1_analoaf(sK0) [avatar definition]
182. l1_analoaf(sK0) <- (6) [avatar component clause 180] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
183. 6 [avatar split clause 112,180]
185. 7 <=> v2_analoaf(sK0) [avatar definition]
187. v2_analoaf(sK0) <- (7) [avatar component clause 185] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
188. 7 [avatar split clause 111,185]
190. 8 <=> v7_struct_0(sK0) [avatar definition]
192. ~v7_struct_0(sK0) <- (~8) [avatar component clause 190] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
193. ~8 [avatar split clause 110,190]
224. l1_struct_0(sK0) <- (6) [resolution 118,182] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
227. 15 <=> l1_struct_0(sK0) [avatar definition]
229. l1_struct_0(sK0) <- (15) [avatar component clause 227] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
230. 15 | ~6 [avatar split clause 224,180,227]
289. 24 <=> v2_struct_0(sK0) [avatar definition]
290. ~v2_struct_0(sK0) <- (~24) [avatar component clause 289] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
291. v2_struct_0(sK0) <- (24) [avatar component clause 289] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
317. r3_diraf(X2,X3,X4,X4) | ~m1_subset_1(X4,u1_struct_0(X2)) | ~m1_subset_1(X4,u1_struct_0(X2)) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l1_analoaf(X2) | v2_struct_0(X2) | ~m1_subset_1(X4,u1_struct_0(X2)) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l1_analoaf(X2) | ~v2_analoaf(X2) | v7_struct_0(X2) [resolution 136,139] {a:1,w:35,wCS:14000,sine:1,thAx:0,allAx:2,thDist:-2}
319. r3_diraf(X9,X10,X11,X10) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~m1_subset_1(X11,u1_struct_0(X9)) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~l1_analoaf(X9) | v2_struct_0(X9) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~m1_subset_1(X11,u1_struct_0(X9)) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~l1_analoaf(X9) | ~v2_analoaf(X9) | v7_struct_0(X9) [resolution 136,140] {a:1,w:39,wCS:15600,sine:1,thAx:0,allAx:2,thDist:-2}
320. r3_diraf(X12,X13,X13,X14) | ~m1_subset_1(X14,u1_struct_0(X12)) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~l1_analoaf(X12) | v2_struct_0(X12) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~m1_subset_1(X14,u1_struct_0(X12)) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~l1_analoaf(X12) | ~v2_analoaf(X12) | v7_struct_0(X12) [resolution 136,141] {a:1,w:39,wCS:15600,sine:1,thAx:0,allAx:2,thDist:-2}
321. r3_diraf(X12,X13,X13,X14) | ~m1_subset_1(X14,u1_struct_0(X12)) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~l1_analoaf(X12) | v2_struct_0(X12) | ~v2_analoaf(X12) | v7_struct_0(X12) [duplicate literal removal 320] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
322. r3_diraf(X9,X10,X11,X10) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~m1_subset_1(X11,u1_struct_0(X9)) | ~l1_analoaf(X9) | v2_struct_0(X9) | ~v2_analoaf(X9) | v7_struct_0(X9) [duplicate literal removal 319] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
324. r3_diraf(X2,X3,X4,X4) | ~m1_subset_1(X4,u1_struct_0(X2)) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l1_analoaf(X2) | v2_struct_0(X2) | ~v2_analoaf(X2) | v7_struct_0(X2) [duplicate literal removal 317] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
342. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1) [resolution 321,159] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
344. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 4) [subsumption resolution 342,172] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
345. ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 4, 5) [subsumption resolution 344,177] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
346. v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 4, 5, 6) [subsumption resolution 345,182] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
353. v2_struct_0(sK0) | v7_struct_0(sK0) <- (~1, 4, 5, 6, 7) [subsumption resolution 346,187] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
354. v2_struct_0(sK0) <- (~1, 4, 5, 6, 7, ~8) [subsumption resolution 353,192] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
355. 24 | 1 | ~4 | ~5 | ~6 | ~7 | 8 [avatar split clause 354,190,185,180,175,170,157,289]
360. v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (24) [resolution 291,122] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
363. ~l1_struct_0(sK0) <- (~8, 24) [subsumption resolution 360,192] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
364. $false <- (~8, 15, 24) [subsumption resolution 363,229] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
365. 8 | ~15 | ~24 [avatar contradiction clause 364]
375. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2) [resolution 324,163] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
377. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 4) [subsumption resolution 375,172] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
378. ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 4, 5) [subsumption resolution 377,177] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
379. v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 4, 5, 6) [subsumption resolution 378,182] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
380. ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 4, 5, 6, ~24) [subsumption resolution 379,290] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
381. v7_struct_0(sK0) <- (~2, 4, 5, 6, 7, ~24) [subsumption resolution 380,187] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
382. $false <- (~2, 4, 5, 6, 7, ~8, ~24) [subsumption resolution 381,192] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
383. 2 | ~4 | ~5 | ~6 | ~7 | 8 | 24 [avatar contradiction clause 382]
385. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~3) [resolution 167,322] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
387. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~3, 5) [subsumption resolution 385,177] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
388. ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~3, 4, 5) [subsumption resolution 387,172] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
389. v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~3, 4, 5, 6) [subsumption resolution 388,182] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
390. ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~3, 4, 5, 6, ~24) [subsumption resolution 389,290] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
391. v7_struct_0(sK0) <- (~3, 4, 5, 6, 7, ~24) [subsumption resolution 390,187] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
392. $false <- (~3, 4, 5, 6, 7, ~8, ~24) [subsumption resolution 391,192] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
393. 3 | ~4 | ~5 | ~6 | ~7 | 8 | 24 [avatar contradiction clause 392]
394. $false [avatar sat refutation 168,173,178,183,188,193,230,355,365,383,393] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:298,thDist:-298}
% SZS output end Proof for diraf__t31_diraf
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 39
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 46
% Generated clauses: 184
% Active clauses: 69
% Passive clauses: 76
% Final active clauses: 60
% Main loop iterations started: 74
% 
% >>> Simplifying Inferences
% Duplicate literals: 39
% Fw subsumption resolutions: 26
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 21
% 
% >>> Generating Inferences
% Binary resolution: 49
% Unit resulting resolution: 4
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 11
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 36
% SAT solver unit clauses: 11
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

