i: [1,35,0,2,1,1,0,0] t125_relat_1
a: [3,65,0,2,1,35]
s: 65
i: [1,34,0,2,1,1,0,0] t125_relat_1
a: [3,70,0,2,1,34]
s: 70
i: [1,37,0,3,1,1,0,0] t125_relat_1
a: [3,55,0,3,1,37]
s: 55
i: [1,36,0,3,1,1,0,0] t125_relat_1
a: [3,60,0,3,1,36]
s: 60
i: [1,38,0,7,1,1,0,0] t125_relat_1
a: [3,50,0,7,1,38]
s: 50
i: [1,46,0,3,1,0,0,1] reflexivity_r1_tarski
s: 46
i: [1,45,0,4,1,0,0,255] existence_m1_subset_1
s: 45
i: [1,41,0,7,2,0,0,255] t3_subset
s: 41
d: [2,72,1,5,1,0,40,41,45]
s: 72
i: [1,42,0,7,2,0,0,255] t3_subset
s: 42
i: [1,39,0,8,3,0,0,255] cc2_relat_1
s: 39
d: [2,74,1,6,2,0,40,39,45]
s: 74
i: [1,44,0,9,3,0,0,1] t1_xboole_1
s: 44
d: [2,77,1,6,2,1,40,44,55]
s: 77
i: [1,43,0,12,3,0,0,1] t123_relat_1
s: 43
i: [1,40,0,14,4,0,0,1] t124_relat_1
s: 40
d: [2,81,2,5,1,1,40,77,72]
a: [3,85,2,5,1,81]
s: 85
d: [2,78,1,6,2,1,40,44,60]
s: 78
d: [2,91,2,5,1,1,40,78,72]
a: [3,95,2,5,1,91]
s: 95
d: [2,75,1,7,3,0,40,39,42]
s: 75
d: [2,87,1,15,4,0,40,43,44]
s: 87
d: [2,101,2,4,2,1,40,75,55]
a: [3,114,2,2,1,101]
s: 114
d: [2,88,1,17,5,0,40,40,44]
s: 88
d: [2,105,3,6,2,1,40,75,95]
d: [2,121,3,4,1,2,21,105,65]
a: [3,125,3,4,1,121]
s: 125
d: [2,134,1,15,4,0,21,88,75]
s: 134
d: [2,79,2,8,2,0,40,44,72]
s: 79
d: [2,89,3,8,2,1,40,85,44]
s: 89
d: [2,135,2,12,3,0,40,134,46]
s: 135
r: 40
d: [2,144,4,7,1,1,40,89,72]
a: [3,148,4,7,1,144]
s: 148
d: [2,99,2,13,4,0,40,75,43]
s: 99
d: [2,100,2,15,5,0,40,75,40]
d: [2,106,2,13,4,0,21,100,75]
s: 106
d: [2,97,3,8,2,1,40,95,44]
s: 97
d: [2,130,2,14,3,0,40,87,72]
s: 130
d: [2,159,4,7,1,1,40,97,72]
a: [3,163,4,7,1,159]
s: 163
d: [2,138,2,14,3,0,40,134,72]
s: 138
d: [2,128,2,17,5,0,40,87,43]
d: [2,132,2,15,4,0,18,128]
s: 132
d: [2,171,5,8,2,1,40,163,75]
d: [2,173,5,6,1,2,21,171,65]
a: [3,177,5,6,1,173]
s: 177
d: [2,129,2,19,6,0,40,87,40]
d: [2,131,2,17,5,0,18,129]
d: [2,133,2,15,4,0,21,131,75]
s: 133
d: [2,191,3,0,0,4,52,65,60,55,50,133]
e: 191
d: [2,192,3,8,3,1,40,133,50]
d: [2,199,3,6,2,2,21,192,65]
d: [2,200,3,3,1,3,21,199,55]
d: [2,201,3,0,0,4,21,200,60]
e: 201
f: 201,35,36,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t125_relat_1
% SZS output start Proof for relat_1__t125_relat_1
1. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => ! [X3] : (v1_relat_1(X3) => ((r1_tarski(X0,X1) & r1_tarski(X2,X3)) => r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X3,X1))))) [input t125_relat_1]
2. ~! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => ! [X3] : (v1_relat_1(X3) => ((r1_tarski(X0,X1) & r1_tarski(X2,X3)) => r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X3,X1))))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
9. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X2,X1)))) [input t123_relat_1]
10. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k7_relat_1(X1,X0),k7_relat_1(X2,X0))))) [input t124_relat_1]
11. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
12. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
13. ~! [X0,X1,X2] : (v1_relat_1(X2) => ! [X3] : (v1_relat_1(X3) => ((r1_tarski(X0,X1) & r1_tarski(X2,X3)) => r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X3,X1))))) [flattening 2]
14. ! [X0,X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k7_relat_1(X1,X0),k7_relat_1(X2,X0))))) [flattening 10]
15. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 12]
16. ! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X2,X1)))) [flattening 9]
17. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 11]
19. ? [X0,X1,X2] : (? [X3] : ((~r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X3,X1)) & (r1_tarski(X0,X1) & r1_tarski(X2,X3))) & v1_relat_1(X3)) & v1_relat_1(X2)) [ennf transformation 13]
20. ? [X0,X1,X2] : (? [X3] : (~r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X3,X1)) & r1_tarski(X0,X1) & r1_tarski(X2,X3) & v1_relat_1(X3)) & v1_relat_1(X2)) [flattening 19]
21. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
22. ! [X0,X1] : (! [X2] : ((r1_tarski(k7_relat_1(X1,X0),k7_relat_1(X2,X0)) | ~r1_tarski(X1,X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [ennf transformation 14]
23. ! [X0,X1] : (! [X2] : (r1_tarski(k7_relat_1(X1,X0),k7_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [flattening 22]
24. ! [X0,X1,X2] : ((r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X2,X1)) | ~r1_tarski(X0,X1)) | ~v1_relat_1(X2)) [ennf transformation 16]
25. ! [X0,X1,X2] : (r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X2,X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X2)) [flattening 24]
26. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 17]
27. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 26]
28. ? [X0,X1,X2] : (? [X3] : (~r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X3,X1)) & r1_tarski(X0,X1) & r1_tarski(X2,X3) & v1_relat_1(X3)) & v1_relat_1(X2)) => (? [X3] : (~r1_tarski(k7_relat_1(sK2,sK0),k7_relat_1(X3,sK1)) & r1_tarski(sK0,sK1) & r1_tarski(sK2,X3) & v1_relat_1(X3)) & v1_relat_1(sK2)) [choice axiom]
29. ? [X3] : (~r1_tarski(k7_relat_1(sK2,sK0),k7_relat_1(X3,sK1)) & r1_tarski(sK0,sK1) & r1_tarski(sK2,X3) & v1_relat_1(X3)) => (~r1_tarski(k7_relat_1(sK2,sK0),k7_relat_1(sK3,sK1)) & r1_tarski(sK0,sK1) & r1_tarski(sK2,sK3) & v1_relat_1(sK3)) [choice axiom]
30. (~r1_tarski(k7_relat_1(sK2,sK0),k7_relat_1(sK3,sK1)) & r1_tarski(sK0,sK1) & r1_tarski(sK2,sK3) & v1_relat_1(sK3)) & v1_relat_1(sK2) [skolemisation 20,29,28]
31. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 15]
35. v1_relat_1(sK3) [cnf transformation 30] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
36. r1_tarski(sK2,sK3) [cnf transformation 30] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
37. r1_tarski(sK0,sK1) [cnf transformation 30] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. ~r1_tarski(k7_relat_1(sK2,sK0),k7_relat_1(sK3,sK1)) [cnf transformation 30] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 21] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. r1_tarski(k7_relat_1(X1,X0),k7_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 23] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 31] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. r1_tarski(k7_relat_1(X2,X0),k7_relat_1(X2,X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X2) [cnf transformation 25] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 27] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
48. 1 <=> r1_tarski(k7_relat_1(sK2,sK0),k7_relat_1(sK3,sK1)) [avatar definition]
50. ~r1_tarski(k7_relat_1(sK2,sK0),k7_relat_1(sK3,sK1)) <- (~1) [avatar component clause 48] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
51. ~1 [avatar split clause 38,48]
53. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
55. r1_tarski(sK0,sK1) <- (2) [avatar component clause 53] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
56. 2 [avatar split clause 37,53]
58. 3 <=> r1_tarski(sK2,sK3) [avatar definition]
60. r1_tarski(sK2,sK3) <- (3) [avatar component clause 58] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
61. 3 [avatar split clause 36,58]
63. 4 <=> v1_relat_1(sK3) [avatar definition]
65. v1_relat_1(sK3) <- (4) [avatar component clause 63] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66. 4 [avatar split clause 35,63]
75. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 39,42] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. ~r1_tarski(X3,k7_relat_1(X2,X0)) | ~v1_relat_1(X2) | r1_tarski(X3,k7_relat_1(X2,X1)) | ~r1_tarski(X0,X1) [resolution 43,44] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
129. ~v1_relat_1(X7) | r1_tarski(k7_relat_1(X8,X9),k7_relat_1(X7,X10)) | ~r1_tarski(X9,X10) | ~r1_tarski(X8,X7) | ~v1_relat_1(X7) | ~v1_relat_1(X8) [resolution 87,40] {a:2,w:19,wCS:7600,sine:1,thAx:0,allAx:3,thDist:-3}
131. ~v1_relat_1(X7) | r1_tarski(k7_relat_1(X8,X9),k7_relat_1(X7,X10)) | ~r1_tarski(X9,X10) | ~r1_tarski(X8,X7) | ~v1_relat_1(X8) [duplicate literal removal 129] {a:2,w:17,wCS:6800,sine:1,thAx:0,allAx:3,thDist:-3}
133. r1_tarski(k7_relat_1(X8,X9),k7_relat_1(X7,X10)) | ~v1_relat_1(X7) | ~r1_tarski(X9,X10) | ~r1_tarski(X8,X7) [subsumption resolution 131,75] {a:2,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
192. ~v1_relat_1(sK3) | ~r1_tarski(sK0,sK1) | ~r1_tarski(sK2,sK3) <- (~1) [resolution 133,50] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
199. ~r1_tarski(sK0,sK1) | ~r1_tarski(sK2,sK3) <- (~1, 4) [subsumption resolution 192,65] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
200. ~r1_tarski(sK2,sK3) <- (~1, 2, 4) [subsumption resolution 199,55] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
201. $false <- (~1, 2, 3, 4) [subsumption resolution 200,60] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
202. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 201]
204. $false [avatar sat refutation 51,56,61,66,202] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:133,thDist:-133}
% SZS output end Proof for relat_1__t125_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 111
% Active clauses: 37
% Passive clauses: 58
% Final active clauses: 36
% Final passive clauses: 20
% Main loop iterations started: 38
% 
% >>> Simplifying Inferences
% Duplicate literals: 7
% Fw subsumption resolutions: 15
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 10
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 63
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 5
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% sat proof minimization: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

