i: [1,112,0,2,1,1,0,0] t61_partfun2
a: [3,183,0,2,1,112]
s: 183
i: [1,110,0,2,1,1,0,0] t61_partfun2
a: [3,193,0,2,1,110]
s: 193
i: [1,109,0,2,1,1,0,0] t61_partfun2
a: [3,198,0,2,1,109]
s: 198
i: [1,107,0,2,1,1,0,0] t61_partfun2
a: [3,208,0,2,1,107]
s: 208
i: [1,113,0,3,1,1,0,0] t61_partfun2
a: [3,178,0,3,1,113]
s: 178
i: [1,111,0,3,1,1,0,0] t61_partfun2
a: [3,188,0,3,1,111]
s: 188
i: [1,108,0,3,1,1,0,0] t61_partfun2
a: [3,203,0,3,1,108]
s: 203
i: [1,114,0,4,1,1,0,0] t61_partfun2
a: [3,173,0,4,1,114]
s: 173
i: [1,116,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,213,0,2,1,116]
s: 213
i: [1,157,0,2,1,0,0,3] rc7_funct_1
a: [3,223,0,2,1,157]
s: 223
i: [1,156,0,2,1,0,0,3] rc7_funct_1
a: [3,228,0,2,1,156]
s: 228
i: [1,159,0,2,1,0,0,1] rc1_funct_1
a: [3,233,0,2,1,159]
s: 233
i: [1,158,0,2,1,0,0,1] rc1_funct_1
a: [3,238,0,2,1,158]
s: 238
i: [1,162,0,2,1,0,0,255] rc5_funct_1
a: [3,243,0,2,1,162]
s: 243
i: [1,161,0,2,1,0,0,255] rc5_funct_1
a: [3,248,0,2,1,161]
s: 248
i: [1,160,0,2,1,0,0,255] rc5_funct_1
a: [3,253,0,2,1,160]
s: 253
i: [1,164,0,2,1,0,0,255] rc2_funct_1
a: [3,258,0,2,1,164]
s: 258
i: [1,163,0,2,1,0,0,255] rc2_funct_1
a: [3,263,0,2,1,163]
s: 263
i: [1,115,0,9,1,1,0,0] t61_partfun2
a: [3,168,0,9,1,115]
s: 168
i: [1,151,0,3,1,0,0,1] reflexivity_r1_tarski
s: 151
i: [1,117,0,3,1,0,0,255] d2_xboole_0
a: [3,218,0,3,1,117]
s: 218
r: 213
d: [2,265,0,2,1,2,23,213,218]
a: [3,269,0,2,1,265]
s: 269
i: [1,118,0,4,2,0,0,3] cc7_funct_1
s: 118
d: [2,271,1,2,1,1,40,118,269]
a: [3,275,1,2,1,271]
s: 275
i: [1,119,0,4,2,0,0,2] cc1_funct_1
s: 119
i: [1,150,0,4,1,0,0,2] existence_m1_subset_1
s: 150
d: [2,277,1,2,1,1,40,119,269]
a: [3,281,1,2,1,277]
s: 281
i: [1,120,0,5,2,0,0,255] t6_boole
s: 120
i: [1,147,0,5,2,0,0,1] t7_boole
s: 147
i: [1,141,0,6,2,0,0,1] t1_subset
s: 141
i: [1,152,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 152
d: [2,287,1,3,1,1,40,147,173]
a: [3,291,1,3,1,287]
s: 291
i: [1,121,0,7,3,0,0,3] cc8_funct_1
s: 121
d: [2,293,1,4,1,1,40,141,173]
a: [3,297,1,4,1,293]
s: 297
i: [1,122,0,7,3,0,0,3] cc8_funct_1
s: 122
i: [1,144,0,7,2,0,0,1] t3_subset
s: 144
d: [2,306,2,5,2,1,40,297,121]
a: [3,310,2,3,1,306]
s: 310
i: [1,145,0,7,2,0,0,1] t3_subset
s: 145
d: [2,299,1,4,1,1,40,152,173]
a: [3,303,1,4,1,299]
s: 303
i: [1,146,0,7,3,0,0,1] t8_boole
s: 146
i: [1,123,0,8,3,0,0,2] cc9_funct_1
s: 123
d: [2,305,1,5,2,0,40,121,150]
s: 305
i: [1,127,0,8,4,0,0,255] cc6_funct_1
s: 127
d: [2,316,1,5,2,0,40,122,150]
s: 316
i: [1,132,0,8,4,0,0,255] cc4_funct_1
s: 132
d: [2,329,0,6,3,0,21,132,119]
s: 329
d: [2,330,1,4,2,1,40,329,269]
a: [3,338,1,2,1,330]
s: 338
i: [1,133,0,8,4,0,0,255] cc5_funct_1
s: 133
d: [2,340,2,4,2,1,44,338,120]
s: 340
r: 329
i: [1,142,0,8,3,0,0,1] t2_subset
s: 142
i: [1,124,0,9,4,0,0,255] fc18_funct_1
s: 124
d: [2,323,1,5,1,0,40,144,150]
s: 323
i: [1,149,0,9,3,0,0,2] t5_subset
s: 149
d: [2,327,1,6,2,0,40,123,150]
s: 327
i: [1,136,0,10,4,0,0,2] cc3_funct_1
s: 136
i: [1,148,0,10,3,0,0,2] t4_subset
s: 148
d: [2,342,1,6,2,0,40,142,150]
s: 342
i: [1,153,0,13,5,0,0,2] fc17_funct_1
s: 153
d: [2,351,2,6,2,0,40,342,152]
s: 351
i: [1,154,0,13,5,0,0,2] fc17_funct_1
s: 154
i: [1,155,0,13,4,0,0,1] dt_k7_partfun1
s: 155
d: [2,283,1,7,3,0,44,120,120]
s: 283
i: [1,137,0,16,6,0,0,2] t2_grfunc_1
s: 137
d: [2,328,1,7,3,0,40,123,145]
s: 328
i: [1,143,0,19,5,0,0,1] d6_partfun1
s: 143
i: [1,138,0,22,7,0,0,2] t2_grfunc_1
s: 138
d: [2,364,2,4,2,1,40,328,178]
a: [3,373,2,2,1,364]
s: 373
i: [1,139,0,22,7,0,0,2] t2_grfunc_1
s: 139
d: [2,375,1,19,5,1,44,168,143]
d: [2,378,1,17,4,2,21,375,183]
d: [2,379,1,14,3,3,21,378,188]
d: [2,380,1,12,2,4,21,379,193]
a: [3,384,1,4,1,380]
s: 384
i: [1,140,0,27,7,0,0,2] t2_grfunc_1
s: 140
d: [2,345,1,7,2,0,40,149,150]
s: 345
d: [2,420,2,6,2,0,40,345,342]
s: 420
d: [2,325,1,8,3,0,40,145,122]
s: 325
d: [2,421,3,6,2,1,40,420,340]
s: 421
d: [2,326,1,8,3,0,40,145,121]
s: 326
d: [2,346,1,8,3,0,40,149,145]
s: 346
d: [2,425,2,6,2,1,40,346,173]
s: 425
d: [2,347,1,8,3,0,40,136,150]
s: 347
d: [2,428,3,14,6,1,40,425,137]
d: [2,430,3,12,5,2,21,428,198]
d: [2,431,3,10,4,3,21,430,208]
s: 431
d: [2,349,1,8,2,0,40,148,150]
s: 349
d: [2,348,1,9,4,0,40,136,145]
s: 348
d: [2,424,4,5,2,1,40,421,305]
s: 424
d: [2,350,1,9,3,0,40,148,145]
s: 350
d: [2,438,2,7,2,1,40,350,173]
s: 438
d: [2,344,1,10,3,0,40,142,145]
s: 344
d: [2,355,1,14,5,0,40,155,122]
s: 355
d: [2,441,3,15,6,1,40,438,137]
d: [2,443,3,13,5,2,21,441,198]
d: [2,444,3,11,4,3,21,443,208]
s: 444
d: [2,356,1,14,5,0,40,155,121]
s: 356
d: [2,453,4,12,5,3,40,444,122]
a: [3,458,4,10,4,453]
s: 458
d: [2,354,1,15,5,0,40,155,142]
s: 354
d: [2,361,1,15,4,0,40,155,144]
s: 361
d: [2,452,4,14,5,3,40,444,142]
d: [2,455,4,11,4,3,21,452,431]
s: 455
d: [2,476,5,0,0,7,52,183,193,178,384,455]
e: 476
d: [2,477,5,7,3,4,40,455,384]
d: [2,484,5,5,2,5,21,477,183]
d: [2,485,5,3,1,6,21,484,193]
d: [2,486,5,0,0,7,21,485,178]
e: 486
r: 384
a: [3,383,1,4,1,380]
s: 383
d: [2,492,2,3,1,1,40,383,147]
a: [3,506,2,3,1,492]
s: 506
a: [3,388,1,8,1,380]
s: 388
d: [2,360,1,16,5,0,40,155,123]
s: 360
d: [2,490,2,4,1,1,40,383,152]
a: [3,496,2,4,1,490]
s: 496
d: [2,376,1,23,8,0,44,155,143]
d: [2,377,1,16,5,0,18,376]
s: 377
d: [2,491,2,4,1,1,40,383,141]
a: [3,501,2,4,1,491]
s: 501
d: [2,359,1,17,5,0,40,155,149]
s: 359
d: [2,357,1,18,5,0,40,155,148]
s: 357
d: [2,535,3,5,2,1,40,501,122]
a: [3,540,3,3,1,535]
s: 540
d: [2,358,1,18,6,0,40,155,136]
s: 358
d: [2,489,2,6,2,1,40,383,346]
s: 489
d: [2,408,1,19,7,0,40,139,147]
s: 408
d: [2,406,1,22,7,0,40,139,152]
s: 406
d: [2,488,2,7,2,1,40,383,350]
s: 488
d: [2,407,1,22,7,0,40,139,141]
s: 407
d: [2,508,2,18,5,1,44,388,143]
d: [2,509,2,14,4,2,21,508,173]
d: [2,510,2,12,3,3,21,509,198]
d: [2,511,2,9,2,4,21,510,203]
d: [2,512,2,7,1,5,21,511,208]
a: [3,516,2,7,1,512]
s: 516
d: [2,568,3,0,0,7,52,208,198,193,183,178,173,516,138]
e: 568
f: 568,512,486,380,107,108,109,110,111,112,113,114,115
% Refutation found. Thanks to Tanya!
% SZS status Theorem for partfun2__t61_partfun2
% SZS output start Proof for partfun2__t61_partfun2
1. ! [X0] : ! [X1] : ((v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v5_relat_1(X2,X0) & v1_relat_1(X2)) => (r1_tarski(X1,X2) => ! [X3] : (r2_hidden(X3,k9_xtuple_0(X1)) => k7_partfun1(X0,X1,X3) = k7_partfun1(X0,X2,X3))))) [input t61_partfun2]
2. ~! [X0] : ! [X1] : ((v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v5_relat_1(X2,X0) & v1_relat_1(X2)) => (r1_tarski(X1,X2) => ! [X3] : (r2_hidden(X3,k9_xtuple_0(X1)) => k7_partfun1(X0,X1,X3) = k7_partfun1(X0,X2,X3))))) [negated conjecture 1]
14. ! [X0] : ! [X1] : ((v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (r2_hidden(X2,k9_xtuple_0(X1)) => k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2))) [input d6_partfun1]
31. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) <=> (! [X2] : (r2_hidden(X2,k9_xtuple_0(X0)) => k1_funct_1(X1,X2) = k1_funct_1(X0,X2)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)))))) [input t2_grfunc_1]
32. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
33. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
34. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
35. ! [X0] : ! [X1] : ! [X2] : ~(v1_xboole_0(X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) [input t5_subset]
39. ~! [X0,X1] : ((v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v5_relat_1(X2,X0) & v1_relat_1(X2)) => (r1_tarski(X1,X2) => ! [X3] : (r2_hidden(X3,k9_xtuple_0(X1)) => k7_partfun1(X0,X1,X3) = k7_partfun1(X0,X2,X3))))) [flattening 2]
41. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 32]
42. ! [X0,X1] : ((v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (r2_hidden(X2,k9_xtuple_0(X1)) => k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2))) [flattening 14]
43. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 33]
46. ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 34]
47. ! [X0,X1,X2] : ~(v1_xboole_0(X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) [flattening 35]
51. ? [X0,X1] : (? [X2] : ((? [X3] : (k7_partfun1(X0,X1,X3) != k7_partfun1(X0,X2,X3) & r2_hidden(X3,k9_xtuple_0(X1))) & r1_tarski(X1,X2)) & (v1_funct_1(X2) & v5_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 39]
52. ? [X0,X1] : (? [X2] : (? [X3] : (k7_partfun1(X0,X1,X3) != k7_partfun1(X0,X2,X3) & r2_hidden(X3,k9_xtuple_0(X1))) & r1_tarski(X1,X2) & v1_funct_1(X2) & v5_relat_1(X2,X0) & v1_relat_1(X2)) & v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 51]
70. ! [X0] : (! [X1] : ((r1_tarski(X0,X1) <=> (! [X2] : (k1_funct_1(X1,X2) = k1_funct_1(X0,X2) | ~r2_hidden(X2,k9_xtuple_0(X0))) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 31]
71. ! [X0] : (! [X1] : ((r1_tarski(X0,X1) <=> (! [X2] : (k1_funct_1(X1,X2) = k1_funct_1(X0,X2) | ~r2_hidden(X2,k9_xtuple_0(X0))) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 70]
73. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 41]
74. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 73]
75. ! [X0,X1] : (! [X2] : (k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2) | ~r2_hidden(X2,k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 42]
76. ! [X0,X1] : (! [X2] : (k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2) | ~r2_hidden(X2,k9_xtuple_0(X1))) | ~v1_funct_1(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 75]
79. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))) [ennf transformation 46]
80. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [flattening 79]
81. ! [X0,X1,X2] : (~v1_xboole_0(X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [ennf transformation 47]
87. ? [X0,X1] : (? [X2] : (? [X3] : (k7_partfun1(X0,X1,X3) != k7_partfun1(X0,X2,X3) & r2_hidden(X3,k9_xtuple_0(X1))) & r1_tarski(X1,X2) & v1_funct_1(X2) & v5_relat_1(X2,X0) & v1_relat_1(X2)) & v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (k7_partfun1(sK0,sK1,X3) != k7_partfun1(sK0,X2,X3) & r2_hidden(X3,k9_xtuple_0(sK1))) & r1_tarski(sK1,X2) & v1_funct_1(X2) & v5_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_funct_1(sK1) & v5_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
88. ? [X2] : (? [X3] : (k7_partfun1(sK0,sK1,X3) != k7_partfun1(sK0,X2,X3) & r2_hidden(X3,k9_xtuple_0(sK1))) & r1_tarski(sK1,X2) & v1_funct_1(X2) & v5_relat_1(X2,sK0) & v1_relat_1(X2)) => (? [X3] : (k7_partfun1(sK0,sK1,X3) != k7_partfun1(sK0,sK2,X3) & r2_hidden(X3,k9_xtuple_0(sK1))) & r1_tarski(sK1,sK2) & v1_funct_1(sK2) & v5_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
89. ? [X3] : (k7_partfun1(sK0,sK1,X3) != k7_partfun1(sK0,sK2,X3) & r2_hidden(X3,k9_xtuple_0(sK1))) => (k7_partfun1(sK0,sK1,sK3) != k7_partfun1(sK0,sK2,sK3) & r2_hidden(sK3,k9_xtuple_0(sK1))) [choice axiom]
90. ((k7_partfun1(sK0,sK1,sK3) != k7_partfun1(sK0,sK2,sK3) & r2_hidden(sK3,k9_xtuple_0(sK1))) & r1_tarski(sK1,sK2) & v1_funct_1(sK2) & v5_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_funct_1(sK1) & v5_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 52,89,88,87]
91. ! [X0] : (! [X1] : (((r1_tarski(X0,X1) | (? [X2] : (k1_funct_1(X1,X2) != k1_funct_1(X0,X2) & r2_hidden(X2,k9_xtuple_0(X0))) | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)))) & ((! [X2] : (k1_funct_1(X1,X2) = k1_funct_1(X0,X2) | ~r2_hidden(X2,k9_xtuple_0(X0))) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 71]
92. ! [X0] : (! [X1] : (((r1_tarski(X0,X1) | ? [X2] : (k1_funct_1(X1,X2) != k1_funct_1(X0,X2) & r2_hidden(X2,k9_xtuple_0(X0))) | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) & ((! [X2] : (k1_funct_1(X1,X2) = k1_funct_1(X0,X2) | ~r2_hidden(X2,k9_xtuple_0(X0))) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 91]
93. ! [X0] : (! [X1] : (((r1_tarski(X0,X1) | ? [X2] : (k1_funct_1(X1,X2) != k1_funct_1(X0,X2) & r2_hidden(X2,k9_xtuple_0(X0))) | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) & ((! [X3] : (k1_funct_1(X1,X3) = k1_funct_1(X0,X3) | ~r2_hidden(X3,k9_xtuple_0(X0))) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [rectify 92]
94. ! [X1,X0] : (? [X2] : (k1_funct_1(X1,X2) != k1_funct_1(X0,X2) & r2_hidden(X2,k9_xtuple_0(X0))) => (k1_funct_1(X1,sK4(X0,X1)) != k1_funct_1(X0,sK4(X0,X1)) & r2_hidden(sK4(X0,X1),k9_xtuple_0(X0)))) [choice axiom]
95. ! [X0] : (! [X1] : (((r1_tarski(X0,X1) | (k1_funct_1(X1,sK4(X0,X1)) != k1_funct_1(X0,sK4(X0,X1)) & r2_hidden(sK4(X0,X1),k9_xtuple_0(X0))) | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) & ((! [X3] : (k1_funct_1(X1,X3) = k1_funct_1(X0,X3) | ~r2_hidden(X3,k9_xtuple_0(X0))) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [skolemisation 93,94]
96. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 43]
107. v1_relat_1(sK1) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
108. v5_relat_1(sK1,sK0) [cnf transformation 90] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
109. v1_funct_1(sK1) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
110. v1_relat_1(sK2) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
111. v5_relat_1(sK2,sK0) [cnf transformation 90] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
112. v1_funct_1(sK2) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
113. r1_tarski(sK1,sK2) [cnf transformation 90] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
114. r2_hidden(sK3,k9_xtuple_0(sK1)) [cnf transformation 90] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
115. k7_partfun1(sK0,sK1,sK3) != k7_partfun1(sK0,sK2,sK3) [cnf transformation 90] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
137. r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 95] {a:0,w:16,wCS:6400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
138. k1_funct_1(X1,X3) = k1_funct_1(X0,X3) | ~r2_hidden(X3,k9_xtuple_0(X0)) | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 95] {a:0,w:22,wCS:8800,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
142. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 74] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
143. k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2) | ~r2_hidden(X2,k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 76] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
145. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 96] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
148. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1) [cnf transformation 80] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
149. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~v1_xboole_0(X2) | ~r2_hidden(X0,X1) [cnf transformation 81] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
166. 1 <=> k7_partfun1(sK0,sK1,sK3) = k7_partfun1(sK0,sK2,sK3) [avatar definition]
168. k7_partfun1(sK0,sK1,sK3) != k7_partfun1(sK0,sK2,sK3) <- (~1) [avatar component clause 166] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
169. ~1 [avatar split clause 115,166]
171. 2 <=> r2_hidden(sK3,k9_xtuple_0(sK1)) [avatar definition]
173. r2_hidden(sK3,k9_xtuple_0(sK1)) <- (2) [avatar component clause 171] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
174. 2 [avatar split clause 114,171]
176. 3 <=> r1_tarski(sK1,sK2) [avatar definition]
178. r1_tarski(sK1,sK2) <- (3) [avatar component clause 176] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
179. 3 [avatar split clause 113,176]
181. 4 <=> v1_funct_1(sK2) [avatar definition]
183. v1_funct_1(sK2) <- (4) [avatar component clause 181] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
184. 4 [avatar split clause 112,181]
186. 5 <=> v5_relat_1(sK2,sK0) [avatar definition]
188. v5_relat_1(sK2,sK0) <- (5) [avatar component clause 186] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
189. 5 [avatar split clause 111,186]
191. 6 <=> v1_relat_1(sK2) [avatar definition]
193. v1_relat_1(sK2) <- (6) [avatar component clause 191] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
194. 6 [avatar split clause 110,191]
196. 7 <=> v1_funct_1(sK1) [avatar definition]
198. v1_funct_1(sK1) <- (7) [avatar component clause 196] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
199. 7 [avatar split clause 109,196]
201. 8 <=> v5_relat_1(sK1,sK0) [avatar definition]
203. v5_relat_1(sK1,sK0) <- (8) [avatar component clause 201] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
204. 8 [avatar split clause 108,201]
206. 9 <=> v1_relat_1(sK1) [avatar definition]
208. v1_relat_1(sK1) <- (9) [avatar component clause 206] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
209. 9 [avatar split clause 107,206]
346. ~r2_hidden(X3,X4) | ~v1_xboole_0(X2) | ~r1_tarski(X4,X2) [resolution 149,145] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
350. ~r2_hidden(X2,X4) | m1_subset_1(X2,X3) | ~r1_tarski(X4,X3) [resolution 148,145] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
375. k7_partfun1(sK0,sK1,sK3) != k1_funct_1(sK2,sK3) | ~r2_hidden(sK3,k9_xtuple_0(sK2)) | ~v1_funct_1(sK2) | ~v5_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1) [superposition 168,143] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
378. k7_partfun1(sK0,sK1,sK3) != k1_funct_1(sK2,sK3) | ~r2_hidden(sK3,k9_xtuple_0(sK2)) | ~v5_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 4) [subsumption resolution 375,183] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
379. k7_partfun1(sK0,sK1,sK3) != k1_funct_1(sK2,sK3) | ~r2_hidden(sK3,k9_xtuple_0(sK2)) | ~v1_relat_1(sK2) <- (~1, 4, 5) [subsumption resolution 378,188] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
380. k7_partfun1(sK0,sK1,sK3) != k1_funct_1(sK2,sK3) | ~r2_hidden(sK3,k9_xtuple_0(sK2)) <- (~1, 4, 5, 6) [subsumption resolution 379,193] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
382. 34 <=> r2_hidden(sK3,k9_xtuple_0(sK2)) [avatar definition]
384. ~r2_hidden(sK3,k9_xtuple_0(sK2)) <- (~34) [avatar component clause 382] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
386. 35 <=> k7_partfun1(sK0,sK1,sK3) = k1_funct_1(sK2,sK3) [avatar definition]
388. k7_partfun1(sK0,sK1,sK3) != k1_funct_1(sK2,sK3) <- (~35) [avatar component clause 386] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
389. ~34 | ~35 | 1 | ~4 | ~5 | ~6 [avatar split clause 380,191,186,181,166,386,382]
425. ~r1_tarski(k9_xtuple_0(sK1),X0) | ~v1_xboole_0(X0) <- (2) [resolution 346,173] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
428. ~v1_xboole_0(k9_xtuple_0(X0)) | ~r1_tarski(sK1,X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2) [resolution 425,137] {a:3,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
430. ~v1_xboole_0(k9_xtuple_0(X0)) | ~r1_tarski(sK1,X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_relat_1(sK1) <- (2, 7) [subsumption resolution 428,198] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
431. ~v1_xboole_0(k9_xtuple_0(X0)) | ~r1_tarski(sK1,X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2, 7, 9) [subsumption resolution 430,208] {a:3,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
438. ~r1_tarski(k9_xtuple_0(sK1),X0) | m1_subset_1(sK3,X0) <- (2) [resolution 350,173] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
441. m1_subset_1(sK3,k9_xtuple_0(X0)) | ~r1_tarski(sK1,X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2) [resolution 438,137] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
443. m1_subset_1(sK3,k9_xtuple_0(X0)) | ~r1_tarski(sK1,X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_relat_1(sK1) <- (2, 7) [subsumption resolution 441,198] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
444. m1_subset_1(sK3,k9_xtuple_0(X0)) | ~r1_tarski(sK1,X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2, 7, 9) [subsumption resolution 443,208] {a:3,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
452. ~r1_tarski(sK1,X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | v1_xboole_0(k9_xtuple_0(X0)) | r2_hidden(sK3,k9_xtuple_0(X0)) <- (2, 7, 9) [resolution 444,142] {a:4,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
455. r2_hidden(sK3,k9_xtuple_0(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~r1_tarski(sK1,X0) <- (2, 7, 9) [subsumption resolution 452,431] {a:4,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
477. ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~r1_tarski(sK1,sK2) <- (2, 7, 9, ~34) [resolution 455,384] {a:5,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
484. ~v1_relat_1(sK2) | ~r1_tarski(sK1,sK2) <- (2, 4, 7, 9, ~34) [subsumption resolution 477,183] {a:5,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
485. ~r1_tarski(sK1,sK2) <- (2, 4, 6, 7, 9, ~34) [subsumption resolution 484,193] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
486. $false <- (2, 3, 4, 6, 7, 9, ~34) [subsumption resolution 485,178] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
487. ~2 | ~3 | ~4 | ~6 | ~7 | ~9 | 34 [avatar contradiction clause 486]
508. k1_funct_1(sK2,sK3) != k1_funct_1(sK1,sK3) | ~r2_hidden(sK3,k9_xtuple_0(sK1)) | ~v1_funct_1(sK1) | ~v5_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~35) [superposition 388,143] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
509. k1_funct_1(sK2,sK3) != k1_funct_1(sK1,sK3) | ~v1_funct_1(sK1) | ~v5_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~35) [subsumption resolution 508,173] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
510. k1_funct_1(sK2,sK3) != k1_funct_1(sK1,sK3) | ~v5_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 7, ~35) [subsumption resolution 509,198] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
511. k1_funct_1(sK2,sK3) != k1_funct_1(sK1,sK3) | ~v1_relat_1(sK1) <- (2, 7, 8, ~35) [subsumption resolution 510,203] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
512. k1_funct_1(sK2,sK3) != k1_funct_1(sK1,sK3) <- (2, 7, 8, 9, ~35) [subsumption resolution 511,208] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
514. 40 <=> k1_funct_1(sK2,sK3) = k1_funct_1(sK1,sK3) [avatar definition]
516. k1_funct_1(sK2,sK3) != k1_funct_1(sK1,sK3) <- (~40) [avatar component clause 514] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
517. ~40 | ~2 | ~7 | ~8 | ~9 | 35 [avatar split clause 512,386,206,201,196,171,514]
568. $false <- (2, 3, 4, 6, 7, 9, ~40) [unit resulting resolution 208,198,193,183,178,173,516,138] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:32,thDist:-32}
571. ~2 | ~3 | ~4 | ~6 | ~7 | ~9 | 40 [avatar contradiction clause 568]
577. $false [avatar sat refutation 169,174,179,184,189,194,199,204,209,389,487,517,571] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:308,thDist:-308}
% SZS output end Proof for partfun2__t61_partfun2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 37
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 58
% Generated clauses: 306
% Active clauses: 106
% Passive clauses: 164
% Final active clauses: 103
% Final passive clauses: 55
% Main loop iterations started: 109
% 
% >>> Simplifying Inferences
% Duplicate literals: 65
% Fw subsumption resolutions: 29
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 16
% Forward subsumptions: 27
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 116
% Unit resulting resolution: 2
% Forward superposition: 18
% Backward superposition: 9
% Self superposition: 5
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 18
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 44
% SAT solver unit clauses: 20
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 9722
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% non unit clause index maintenance: 0.001 s
% backward subsumption index maintenance: 0.001 s
% backward superposition index maintenance: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

