i: [1,86,0,2,1,1,0,0] t21_mcart_1
a: [3,135,0,2,1,86]
s: 135
i: [1,87,0,3,1,1,0,0] t21_mcart_1
a: [3,130,0,3,1,87]
s: 130
i: [1,88,0,7,1,1,0,0] t21_mcart_1
a: [3,125,0,7,1,88]
s: 125
i: [1,89,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,140,0,2,1,89]
s: 140
i: [1,90,0,2,1,0,0,255] fc1_xboole_0
a: [3,145,0,2,1,90]
s: 145
i: [1,116,0,2,1,0,0,2] rc2_xboole_0
a: [3,155,0,2,1,116]
s: 155
i: [1,117,0,2,1,0,0,2] rc1_xtuple_0
a: [3,160,0,2,1,117]
s: 160
i: [1,118,0,2,1,0,0,2] rc1_xboole_0
a: [3,165,0,2,1,118]
s: 165
i: [1,120,0,2,1,0,0,1] rc1_relat_1
a: [3,170,0,2,1,120]
s: 170
i: [1,119,0,2,1,0,0,1] rc1_relat_1
a: [3,175,0,2,1,119]
s: 175
i: [1,121,0,2,1,0,0,255] rc2_relat_1
a: [3,180,0,2,1,121]
s: 180
i: [1,92,0,3,1,0,0,255] fc2_xboole_0
s: 92
i: [1,91,0,3,1,0,0,255] d2_xboole_0
a: [3,150,0,3,1,91]
s: 150
r: 140
i: [1,95,0,4,2,0,0,1] cc1_relat_1
s: 95
i: [1,107,0,4,1,0,0,255] existence_m1_subset_1
s: 107
i: [1,108,0,4,1,0,0,255] fc3_xboole_0
s: 108
d: [2,183,1,2,1,1,40,95,165]
a: [3,188,1,2,1,183]
s: 188
i: [1,109,0,4,1,0,0,2] fc1_xtuple_0
s: 109
d: [2,184,1,2,1,1,40,95,145]
a: [3,193,1,2,1,184]
s: 193
i: [1,96,0,5,2,0,0,255] t6_boole
s: 96
i: [1,106,0,5,2,0,0,1] t7_boole
s: 106
d: [2,198,1,2,1,1,40,106,130]
a: [3,202,1,2,1,198]
s: 202
i: [1,110,0,5,1,0,0,255] fc5_relat_1
s: 110
i: [1,100,0,6,2,0,0,1] d1_relat_1
s: 100
i: [1,104,0,6,2,0,0,255] t1_subset
s: 104
i: [1,111,0,6,1,0,0,1] rd1_xtuple_0
s: 111
d: [2,205,1,3,1,1,40,104,130]
a: [3,210,1,3,1,205]
s: 210
i: [1,112,0,6,1,0,0,1] rd2_xtuple_0
s: 112
i: [1,114,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 114
i: [1,105,0,7,3,0,0,1] t8_boole
s: 105
i: [1,113,0,7,1,0,0,255] commutativity_k2_tarski
s: 113
d: [2,212,1,3,1,1,40,114,130]
a: [3,217,1,3,1,212]
s: 217
i: [1,101,0,8,2,0,0,1] d1_relat_1
s: 101
d: [2,206,1,6,2,0,40,104,100]
s: 206
i: [1,103,0,8,3,0,0,255] t2_subset
s: 103
i: [1,115,0,8,1,0,0,255] fc7_relat_1
s: 115
d: [2,213,1,6,2,0,40,114,100]
s: 213
i: [1,93,0,9,2,0,0,1] t8_mcart_1
s: 93
d: [2,228,1,5,2,1,44,125,93]
d: [2,234,1,2,1,1,19,228]
a: [3,240,1,2,1,234]
s: 240
i: [1,94,0,9,2,0,0,1] rd3_xtuple_0
s: 94
i: [1,102,0,10,1,0,0,255] d5_tarski
s: 102
d: [2,249,1,4,1,0,44,108,102]
s: 249
i: [1,99,0,12,3,0,0,1] d1_relat_1
s: 99
d: [2,250,2,4,2,0,44,249,93]
s: 250
d: [2,231,1,5,2,0,44,110,93]
s: 231
d: [2,221,1,6,2,0,40,103,107]
s: 221
d: [2,259,3,2,1,1,40,250,160]
a: [3,264,3,2,1,259]
s: 264
d: [2,195,1,7,3,0,44,96,96]
s: 195
d: [2,266,2,6,2,0,40,221,114]
s: 266
d: [2,252,1,7,3,0,44,109,99]
s: 252
d: [2,269,2,0,0,3,52,135,240,130,252]
e: 269
d: [2,270,2,4,2,1,40,252,130]
d: [2,274,2,2,1,2,21,270,240]
d: [2,275,2,0,0,3,21,274,135]
e: 275
f: 275,234,86,87,88
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t21_mcart_1
% SZS output start Proof for mcart_1__t21_mcart_1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => (r2_hidden(X0,X1) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0)) [input t21_mcart_1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => (r2_hidden(X0,X1) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0)) [negated conjecture 1]
8. ! [X0] : (v1_relat_1(X0) <=> ! [X1] : ~(! [X2] : ! [X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) [input d1_relat_1]
21. ! [X0,X1] : v1_xtuple_0(k4_tarski(X0,X1)) [input fc1_xtuple_0]
39. ! [X0] : (v1_xtuple_0(X0) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0) [input t8_mcart_1]
40. ~! [X0,X1] : (v1_relat_1(X1) => (r2_hidden(X0,X1) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0)) [flattening 2]
41. ! [X0] : (v1_relat_1(X0) <=> ! [X1] : ~(! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) [flattening 8]
50. ? [X0,X1] : ((k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) != X0 & r2_hidden(X0,X1)) & v1_relat_1(X1)) [ennf transformation 40]
51. ? [X0,X1] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) != X0 & r2_hidden(X0,X1) & v1_relat_1(X1)) [flattening 50]
52. ! [X0] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0)) [ennf transformation 39]
60. ! [X0] : (v1_relat_1(X0) <=> ! [X1] : (? [X2,X3] : k4_tarski(X2,X3) = X1 | ~r2_hidden(X1,X0))) [ennf transformation 41]
67. ? [X0,X1] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) != X0 & r2_hidden(X0,X1) & v1_relat_1(X1)) => (sK0 != k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) & r2_hidden(sK0,sK1) & v1_relat_1(sK1)) [choice axiom]
68. sK0 != k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) & r2_hidden(sK0,sK1) & v1_relat_1(sK1) [skolemisation 51,67]
69. ! [X0] : ((v1_relat_1(X0) | ? [X1] : (! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) & (! [X1] : (? [X2,X3] : k4_tarski(X2,X3) = X1 | ~r2_hidden(X1,X0)) | ~v1_relat_1(X0))) [nnf transformation 60]
70. ! [X0] : ((v1_relat_1(X0) | ? [X1] : (! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) & (! [X4] : (? [X5,X6] : k4_tarski(X5,X6) = X4 | ~r2_hidden(X4,X0)) | ~v1_relat_1(X0))) [rectify 69]
71. ! [X0] : (? [X1] : (! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0)) => (! [X3,X2] : k4_tarski(X2,X3) != sK2(X0) & r2_hidden(sK2(X0),X0))) [choice axiom]
72. ! [X4] : (? [X5,X6] : k4_tarski(X5,X6) = X4 => k4_tarski(sK3(X4),sK4(X4)) = X4) [choice axiom]
73. ! [X0] : ((v1_relat_1(X0) | (! [X2,X3] : k4_tarski(X2,X3) != sK2(X0) & r2_hidden(sK2(X0),X0))) & (! [X4] : (k4_tarski(sK3(X4),sK4(X4)) = X4 | ~r2_hidden(X4,X0)) | ~v1_relat_1(X0))) [skolemisation 70,72,71]
86. v1_relat_1(sK1) [cnf transformation 68] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
87. r2_hidden(sK0,sK1) [cnf transformation 68] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
88. sK0 != k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) [cnf transformation 68] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
93. k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0) [cnf transformation 52] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
99. k4_tarski(sK3(X4),sK4(X4)) = X4 | ~r2_hidden(X4,X0) | ~v1_relat_1(X0) [cnf transformation 73] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
109. v1_xtuple_0(k4_tarski(X0,X1)) [cnf transformation 21] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
123. 1 <=> sK0 = k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) [avatar definition]
125. sK0 != k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (~1) [avatar component clause 123] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
126. ~1 [avatar split clause 88,123]
128. 2 <=> r2_hidden(sK0,sK1) [avatar definition]
130. r2_hidden(sK0,sK1) <- (2) [avatar component clause 128] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
131. 2 [avatar split clause 87,128]
133. 3 <=> v1_relat_1(sK1) [avatar definition]
135. v1_relat_1(sK1) <- (3) [avatar component clause 133] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
136. 3 [avatar split clause 86,133]
228. sK0 != sK0 | ~v1_xtuple_0(sK0) <- (~1) [superposition 125,93] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
234. ~v1_xtuple_0(sK0) <- (~1) [trivial inequality removal 228] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
238. 18 <=> v1_xtuple_0(sK0) [avatar definition]
240. ~v1_xtuple_0(sK0) <- (~18) [avatar component clause 238] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
241. ~18 | 1 [avatar split clause 234,123,238]
252. ~r2_hidden(X3,X4) | v1_xtuple_0(X3) | ~v1_relat_1(X4) [superposition 109,99] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
270. v1_xtuple_0(sK0) | ~v1_relat_1(sK1) <- (2) [resolution 252,130] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
274. ~v1_relat_1(sK1) <- (2, ~18) [subsumption resolution 270,240] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
275. $false <- (2, 3, ~18) [subsumption resolution 274,135] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
276. ~2 | ~3 | 18 [avatar contradiction clause 275]
277. $false [avatar sat refutation 126,131,136,241,276] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:86,thDist:-86}
% SZS output end Proof for mcart_1__t21_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 113
% Active clauses: 48
% Passive clauses: 67
% Final active clauses: 47
% Final passive clauses: 17
% Main loop iterations started: 50
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Forward subsumptions: 13
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Binary resolution: 19
% Unit resulting resolution: 1
% Forward superposition: 9
% Backward superposition: 21
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 12
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% SAT solver time: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

