i: [1,32,0,8,2,1,0,0] t21_gate_1
a: [3,67,0,2,1,32]
s: 67
i: [1,35,0,8,2,1,0,0] t21_gate_1
a: [3,53,0,6,1,35]
s: 53
i: [1,36,0,2,1,0,0,255] fc1_xboole_0
a: [3,74,0,2,1,36]
s: 74
i: [1,44,0,2,1,0,0,1] rc2_xboole_0
a: [3,79,0,2,1,44]
s: 79
i: [1,45,0,2,1,0,0,1] rc1_xboole_0
a: [3,84,0,2,1,45]
s: 84
i: [1,37,0,5,2,0,0,255] fc2_gate_1
s: 37
i: [1,38,0,5,2,0,0,255] fc1_gate_1
s: 38
i: [1,39,0,10,2,0,0,255] d15_gate_1
s: 39
r: 67
a: [3,48,0,2,1,35]
s: 48
a: [3,68,0,2,1,32]
s: 68
i: [1,40,0,10,2,0,0,255] d15_gate_1
s: 40
d: [2,87,1,5,2,1,44,53,39]
d: [2,88,1,3,1,2,21,87,67]
a: [3,92,1,3,1,88]
s: 92
d: [2,97,2,0,0,2,52,74,92,38]
e: 97
d: [2,98,2,2,1,1,40,92,38]
d: [2,100,2,0,0,2,21,98,74]
e: 100
r: 92
i: [1,33,0,8,2,1,0,0] t21_gate_1
a: [3,63,0,2,1,33]
s: 63
a: [3,91,1,3,1,88]
s: 91
i: [1,41,0,10,2,0,0,255] d15_gate_1
s: 41
i: [1,34,0,8,2,1,0,0] t21_gate_1
a: [3,58,0,2,1,34]
s: 58
i: [1,42,0,10,2,0,0,255] d15_gate_1
s: 42
d: [2,106,1,5,2,1,44,53,42]
d: [2,107,1,2,1,2,21,106,91]
d: [2,108,1,0,0,3,21,107,48]
e: 108
r: 48
r: 53
a: [3,49,0,2,1,35]
s: 49
a: [3,52,0,6,1,35]
s: 52
i: [1,43,0,15,5,0,0,255] d15_gate_1
s: 43
d: [2,122,1,10,5,1,44,52,43]
d: [2,123,1,8,4,2,21,122,74]
d: [2,124,1,6,3,3,21,123,49]
d: [2,125,1,4,2,4,21,124,58]
d: [2,126,1,2,1,5,21,125,63]
d: [2,127,1,0,0,6,21,126,68]
e: 127
i: [1,26,0,14,5,1,0,0] t21_gate_1
d: [2,95,1,5,2,1,44,53,40]
d: [2,103,1,5,2,1,44,53,41]
d: [2,104,1,2,1,2,21,103,91]
f: 127,108,104,100,95,87,36,26,32,33,34,35
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t21_gate_1
% SZS output start Proof for gate_1__t21_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (~(v1_xboole_0(k15_gate_1(X0,X1,X2,X3)) & ~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & ~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k15_gate_1(X0,X1,X2,X3)))) [input t21_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (~(v1_xboole_0(k15_gate_1(X0,X1,X2,X3)) & ~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & ~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k15_gate_1(X0,X1,X2,X3)))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (~(k15_gate_1(X0,X1,X2,X3) != k1_xboole_0 & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) & (~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k15_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0))) [input d15_gate_1]
14. ~! [X0,X1,X2,X3] : (~(v1_xboole_0(k15_gate_1(X0,X1,X2,X3)) & ~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & ~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k15_gate_1(X0,X1,X2,X3)))) [flattening 2]
15. ! [X0,X1,X2,X3] : (~(k15_gate_1(X0,X1,X2,X3) != k1_xboole_0 & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) & (~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k15_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0))) [flattening 13]
16. ? [X0,X1,X2,X3] : ((v1_xboole_0(k15_gate_1(X0,X1,X2,X3)) & (~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) | (v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k15_gate_1(X0,X1,X2,X3)))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2,X3] : ((k15_gate_1(X0,X1,X2,X3) = k1_xboole_0 | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) & (k15_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0) | (v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)))) [ennf transformation 15]
20. ? [X0,X1,X2,X3] : ((v1_xboole_0(k15_gate_1(X0,X1,X2,X3)) & (~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) | (v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k15_gate_1(X0,X1,X2,X3)))) => ((v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) & (~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0))) | (v1_xboole_0(sK3) & v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0) & ~v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)))) [choice axiom]
21. (v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) & (~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0))) | (v1_xboole_0(sK3) & v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0) & ~v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3))) [skolemisation 16,20]
26. ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | ~v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 21] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) | v1_xboole_0(sK0) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) | v1_xboole_0(sK1) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) | v1_xboole_0(sK2) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) | v1_xboole_0(sK3) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
38. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. k15_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. k15_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X1) [cnf transformation 19] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
41. k15_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X2) [cnf transformation 19] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. k15_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X3) [cnf transformation 19] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. k15_gate_1(X0,X1,X2,X3) = k1_xboole_0 | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
47. 1 <=> v1_xboole_0(sK3) [avatar definition]
48. ~v1_xboole_0(sK3) <- (~1) [avatar component clause 47] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
49. v1_xboole_0(sK3) <- (1) [avatar component clause 47] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
51. 2 <=> v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) [avatar definition]
52. ~v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) <- (~2) [avatar component clause 51] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
53. v1_xboole_0(k15_gate_1(sK0,sK1,sK2,sK3)) <- (2) [avatar component clause 51] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
54. 1 | 2 [avatar split clause 35,51,47]
56. 3 <=> v1_xboole_0(sK2) [avatar definition]
58. v1_xboole_0(sK2) <- (3) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
59. 3 | 2 [avatar split clause 34,51,56]
61. 4 <=> v1_xboole_0(sK1) [avatar definition]
63. v1_xboole_0(sK1) <- (4) [avatar component clause 61] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
64. 4 | 2 [avatar split clause 33,51,61]
66. 5 <=> v1_xboole_0(sK0) [avatar definition]
68. v1_xboole_0(sK0) <- (5) [avatar component clause 66] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
69. 5 | 2 [avatar split clause 32,51,66]
70. ~2 | ~5 | ~4 | ~3 | ~1 [avatar split clause 26,47,56,61,66,51]
72. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
74. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 72] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
75. 6 [avatar split clause 36,72]
87. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK0) <- (2) [superposition 53,39] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
90. 9 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
91. ~v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (~9) [avatar component clause 90] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (9) [avatar component clause 90] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
94. 5 | 9 | ~2 [avatar split clause 87,51,90,66]
95. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK1) <- (2) [superposition 53,40] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. 4 | 9 | ~2 [avatar split clause 95,51,90,61]
98. ~v1_xboole_0(k1_xboole_0) <- (9) [resolution 92,38] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
100. $false <- (6, 9) [subsumption resolution 98,74] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
101. ~6 | ~9 [avatar contradiction clause 100]
103. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK2) <- (2) [superposition 53,41] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. v1_xboole_0(sK2) <- (2, ~9) [subsumption resolution 103,91] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 3 | ~2 | 9 [avatar split clause 104,90,51,56]
106. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK3) <- (2) [superposition 53,42] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
107. v1_xboole_0(sK3) <- (2, ~9) [subsumption resolution 106,91] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
108. $false <- (~1, 2, ~9) [subsumption resolution 107,48] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
109. 1 | ~2 | 9 [avatar contradiction clause 108]
122. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~2) [superposition 52,43] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
123. ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~2, 6) [subsumption resolution 122,74] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 6) [subsumption resolution 123,49] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
125. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 3, 6) [subsumption resolution 124,58] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. ~v1_xboole_0(sK0) <- (1, ~2, 3, 4, 6) [subsumption resolution 125,63] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
127. $false <- (1, ~2, 3, 4, 5, 6) [subsumption resolution 126,68] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. ~1 | 2 | ~3 | ~4 | ~5 | ~6 [avatar contradiction clause 127]
129. $false [avatar sat refutation 54,59,64,69,70,75,94,96,101,105,109,128] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:70,thDist:-70}
% SZS output end Proof for gate_1__t21_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 20
% Generated clauses: 65
% Active clauses: 20
% Passive clauses: 20
% Final active clauses: 16
% Main loop iterations started: 20
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 1
% Forward superposition: 8
% Backward superposition: 9
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 17
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 3
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 9338
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

