i: [1,71,0,2,1,1,0,0] t5_conmetr1
a: [3,100,0,2,1,71]
s: 100
i: [1,70,0,2,1,1,0,0] t5_conmetr1
a: [3,105,0,2,1,70]
s: 105
i: [1,69,0,2,1,1,0,0] t5_conmetr1
a: [3,110,0,2,1,69]
s: 110
i: [1,68,0,2,1,1,0,0] t5_conmetr1
a: [3,115,0,2,1,68]
s: 115
i: [1,67,0,2,1,1,0,0] t5_conmetr1
a: [3,120,0,2,1,67]
s: 120
i: [1,66,0,2,1,1,0,0] t5_conmetr1
a: [3,125,0,2,1,66]
s: 125
i: [1,95,0,2,1,0,0,1] existence_l1_analoaf
a: [3,130,0,2,1,95]
s: 130
i: [1,96,0,2,1,0,0,2] existence_l1_struct_0
a: [3,135,0,2,1,96]
s: 135
i: [1,72,0,4,2,0,0,1] dt_l1_analoaf
s: 72
i: [1,73,0,6,3,0,0,255] cc7_struct_0
s: 73
i: [1,74,0,6,3,0,0,255] cc5_struct_0
s: 74
d: [2,137,1,2,1,1,40,72,110]
a: [3,142,1,2,1,137]
s: 142
i: [1,76,0,6,3,0,0,255] cc17_struct_0
s: 76
d: [2,138,1,2,1,1,40,72,130]
a: [3,147,1,2,1,138]
s: 147
i: [1,77,0,6,3,0,0,255] cc2_struct_0
s: 77
i: [1,78,0,6,3,0,0,255] cc1_struct_0
s: 78
d: [2,149,1,4,2,1,40,73,135]
a: [3,153,1,2,1,149]
s: 153
i: [1,81,0,6,3,0,0,255] cc4_struct_0
s: 81
d: [2,160,2,4,2,1,40,147,73]
a: [3,164,2,2,1,160]
s: 164
i: [1,82,0,6,3,0,0,255] cc6_struct_0
s: 82
i: [1,79,0,7,3,0,0,255] cc8_struct_0
s: 79
i: [1,83,0,7,3,0,0,255] cc9_struct_0
s: 83
i: [1,84,0,7,3,0,0,255] cc11_struct_0
s: 84
i: [1,85,0,7,3,0,0,255] cc11_struct_0
s: 85
i: [1,86,0,9,4,0,0,255] cc10_struct_0
s: 86
i: [1,87,0,12,6,0,0,2] t4_conmetr1
s: 87
i: [1,88,0,12,6,0,0,2] t3_conmetr1
s: 88
i: [1,89,0,12,6,0,0,1] t2_conmetr1
s: 89
i: [1,90,0,12,6,0,0,1] t2_conmetr1
s: 90
i: [1,92,0,12,6,0,0,1] t1_conmetr1
s: 92
i: [1,93,0,12,6,0,0,1] t1_conmetr1
s: 93
d: [2,176,1,10,5,1,40,92,105]
d: [2,177,1,8,4,2,21,176,110]
d: [2,178,1,6,3,3,21,177,115]
d: [2,179,1,4,2,4,21,178,120]
d: [2,180,1,2,1,5,21,179,125]
a: [3,184,1,2,1,180]
s: 184
i: [1,91,0,14,7,0,0,1] t2_conmetr1
s: 91
d: [2,186,1,10,5,1,40,93,105]
d: [2,187,1,8,4,2,21,186,110]
d: [2,188,1,6,3,3,21,187,115]
d: [2,189,1,4,2,4,21,188,120]
d: [2,190,1,2,1,5,21,189,125]
a: [3,194,1,2,1,190]
s: 194
i: [1,94,0,14,7,0,0,1] t1_conmetr1
s: 94
d: [2,196,2,10,5,1,40,184,88]
d: [2,197,2,8,4,2,21,196,110]
d: [2,198,2,6,3,3,21,197,115]
d: [2,199,2,4,2,4,21,198,120]
d: [2,200,2,2,1,5,21,199,125]
a: [3,204,2,2,1,200]
s: 204
d: [2,206,2,10,5,1,40,194,87]
d: [2,207,2,8,4,2,21,206,110]
d: [2,208,2,6,3,3,21,207,115]
d: [2,209,2,4,2,4,21,208,120]
d: [2,210,2,2,1,5,21,209,125]
a: [3,214,2,2,1,210]
s: 214
d: [2,217,3,0,0,7,52,125,120,115,110,204,100,214,91]
e: 217
d: [2,218,3,12,6,1,40,214,91]
d: [2,220,3,10,5,2,21,218,100]
d: [2,221,3,8,4,3,21,220,204]
d: [2,222,3,6,3,4,21,221,110]
d: [2,223,3,4,2,5,21,222,115]
d: [2,224,3,2,1,6,21,223,120]
d: [2,225,3,0,0,7,21,224,125]
e: 225
f: 225,210,200,190,180,66,67,68,69,70,71
% Refutation found. Thanks to Tanya!
% SZS status Theorem for conmetr1__t5_conmetr1
% SZS output start Proof for conmetr1__t5_conmetr1
1. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v4_conmetr1(X0) => v3_conmetr1(X0))) [input t5_conmetr1]
2. ~! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v4_conmetr1(X0) => v3_conmetr1(X0))) [negated conjecture 1]
19. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v4_conmetr1(X0) <=> (v6_conmetr1(X0) & v5_conmetr1(X0)))) [input t1_conmetr1]
20. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v3_conmetr1(X0) <=> (v2_conmetr1(X0) & v1_conmetr1(X0)))) [input t2_conmetr1]
21. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v5_conmetr1(X0) => v1_conmetr1(X0))) [input t3_conmetr1]
22. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v6_conmetr1(X0) => v2_conmetr1(X0))) [input t4_conmetr1]
23. ? [X0] : ((~v3_conmetr1(X0) & v4_conmetr1(X0)) & (l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
24. ? [X0] : (~v3_conmetr1(X0) & v4_conmetr1(X0) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) [flattening 23]
48. ! [X0] : ((v2_conmetr1(X0) | ~v6_conmetr1(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 22]
49. ! [X0] : (v2_conmetr1(X0) | ~v6_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 48]
50. ! [X0] : ((v1_conmetr1(X0) | ~v5_conmetr1(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 21]
51. ! [X0] : (v1_conmetr1(X0) | ~v5_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 50]
52. ! [X0] : ((v3_conmetr1(X0) <=> (v2_conmetr1(X0) & v1_conmetr1(X0))) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 20]
53. ! [X0] : ((v3_conmetr1(X0) <=> (v2_conmetr1(X0) & v1_conmetr1(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 52]
54. ! [X0] : ((v4_conmetr1(X0) <=> (v6_conmetr1(X0) & v5_conmetr1(X0))) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 19]
55. ! [X0] : ((v4_conmetr1(X0) <=> (v6_conmetr1(X0) & v5_conmetr1(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 54]
56. ? [X0] : (~v3_conmetr1(X0) & v4_conmetr1(X0) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (~v3_conmetr1(sK0) & v4_conmetr1(sK0) & l1_analoaf(sK0) & v2_diraf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
57. ~v3_conmetr1(sK0) & v4_conmetr1(sK0) & l1_analoaf(sK0) & v2_diraf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0) [skolemisation 24,56]
58. ! [X0] : (((v3_conmetr1(X0) | (~v2_conmetr1(X0) | ~v1_conmetr1(X0))) & ((v2_conmetr1(X0) & v1_conmetr1(X0)) | ~v3_conmetr1(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [nnf transformation 53]
59. ! [X0] : (((v3_conmetr1(X0) | ~v2_conmetr1(X0) | ~v1_conmetr1(X0)) & ((v2_conmetr1(X0) & v1_conmetr1(X0)) | ~v3_conmetr1(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 58]
60. ! [X0] : (((v4_conmetr1(X0) | (~v6_conmetr1(X0) | ~v5_conmetr1(X0))) & ((v6_conmetr1(X0) & v5_conmetr1(X0)) | ~v4_conmetr1(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [nnf transformation 55]
61. ! [X0] : (((v4_conmetr1(X0) | ~v6_conmetr1(X0) | ~v5_conmetr1(X0)) & ((v6_conmetr1(X0) & v5_conmetr1(X0)) | ~v4_conmetr1(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 60]
66. ~v7_struct_0(sK0) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
67. v1_diraf(sK0) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
68. v2_diraf(sK0) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. l1_analoaf(sK0) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. v4_conmetr1(sK0) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
71. ~v3_conmetr1(sK0) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
87. ~v6_conmetr1(X0) | v2_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 49] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
88. ~v5_conmetr1(X0) | v1_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 51] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
91. ~v2_conmetr1(X0) | v3_conmetr1(X0) | ~v1_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 59] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
92. ~v4_conmetr1(X0) | v5_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 61] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
93. ~v4_conmetr1(X0) | v6_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 61] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
98. 1 <=> v3_conmetr1(sK0) [avatar definition]
100. ~v3_conmetr1(sK0) <- (~1) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
101. ~1 [avatar split clause 71,98]
103. 2 <=> v4_conmetr1(sK0) [avatar definition]
105. v4_conmetr1(sK0) <- (2) [avatar component clause 103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
106. 2 [avatar split clause 70,103]
108. 3 <=> l1_analoaf(sK0) [avatar definition]
110. l1_analoaf(sK0) <- (3) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
111. 3 [avatar split clause 69,108]
113. 4 <=> v2_diraf(sK0) [avatar definition]
115. v2_diraf(sK0) <- (4) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. 4 [avatar split clause 68,113]
118. 5 <=> v1_diraf(sK0) [avatar definition]
120. v1_diraf(sK0) <- (5) [avatar component clause 118] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
121. 5 [avatar split clause 67,118]
123. 6 <=> v7_struct_0(sK0) [avatar definition]
125. ~v7_struct_0(sK0) <- (~6) [avatar component clause 123] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
126. ~6 [avatar split clause 66,123]
176. v5_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2) [resolution 92,105] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
177. v5_conmetr1(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3) [subsumption resolution 176,110] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
178. v5_conmetr1(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3, 4) [subsumption resolution 177,115] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
179. v5_conmetr1(sK0) | v7_struct_0(sK0) <- (2, 3, 4, 5) [subsumption resolution 178,120] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
180. v5_conmetr1(sK0) <- (2, 3, 4, 5, ~6) [subsumption resolution 179,125] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
182. 15 <=> v5_conmetr1(sK0) [avatar definition]
184. v5_conmetr1(sK0) <- (15) [avatar component clause 182] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
185. 15 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 180,123,118,113,108,103,182]
186. v6_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2) [resolution 93,105] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
187. v6_conmetr1(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3) [subsumption resolution 186,110] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
188. v6_conmetr1(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3, 4) [subsumption resolution 187,115] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
189. v6_conmetr1(sK0) | v7_struct_0(sK0) <- (2, 3, 4, 5) [subsumption resolution 188,120] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
190. v6_conmetr1(sK0) <- (2, 3, 4, 5, ~6) [subsumption resolution 189,125] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
192. 16 <=> v6_conmetr1(sK0) [avatar definition]
194. v6_conmetr1(sK0) <- (16) [avatar component clause 192] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
195. 16 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 190,123,118,113,108,103,192]
196. v1_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (15) [resolution 184,88] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
197. v1_conmetr1(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 15) [subsumption resolution 196,110] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
198. v1_conmetr1(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 4, 15) [subsumption resolution 197,115] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
199. v1_conmetr1(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 15) [subsumption resolution 198,120] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
200. v1_conmetr1(sK0) <- (3, 4, 5, ~6, 15) [subsumption resolution 199,125] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
202. 17 <=> v1_conmetr1(sK0) [avatar definition]
204. v1_conmetr1(sK0) <- (17) [avatar component clause 202] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
205. 17 | ~3 | ~4 | ~5 | 6 | ~15 [avatar split clause 200,182,123,118,113,108,202]
206. v2_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (16) [resolution 194,87] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
207. v2_conmetr1(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 16) [subsumption resolution 206,110] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
208. v2_conmetr1(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 4, 16) [subsumption resolution 207,115] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
209. v2_conmetr1(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 16) [subsumption resolution 208,120] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
210. v2_conmetr1(sK0) <- (3, 4, 5, ~6, 16) [subsumption resolution 209,125] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
212. 18 <=> v2_conmetr1(sK0) [avatar definition]
214. v2_conmetr1(sK0) <- (18) [avatar component clause 212] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
215. 18 | ~3 | ~4 | ~5 | 6 | ~16 [avatar split clause 210,192,123,118,113,108,212]
218. v3_conmetr1(sK0) | ~v1_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (18) [resolution 214,91] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
220. ~v1_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 18) [subsumption resolution 218,100] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
221. ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 17, 18) [subsumption resolution 220,204] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
222. ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 3, 17, 18) [subsumption resolution 221,110] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
223. ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 3, 4, 17, 18) [subsumption resolution 222,115] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
224. v7_struct_0(sK0) <- (~1, 3, 4, 5, 17, 18) [subsumption resolution 223,120] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
225. $false <- (~1, 3, 4, 5, ~6, 17, 18) [subsumption resolution 224,125] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
226. 1 | ~3 | ~4 | ~5 | 6 | ~17 | ~18 [avatar contradiction clause 225]
227. $false [avatar sat refutation 101,106,111,116,121,126,185,195,205,215,226] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:109,thDist:-109}
% SZS output end Proof for conmetr1__t5_conmetr1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 31
% Generated clauses: 87
% Active clauses: 33
% Passive clauses: 37
% Final active clauses: 33
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 22
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 14
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 8
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

