i: [1,47,0,3,1,1,0,0] t75_xboole_1
a: [3,82,0,3,1,47]
s: 82
i: [1,48,0,5,1,1,0,0] t75_xboole_1
a: [3,77,0,5,1,48]
s: 77
i: [1,49,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,87,0,2,1,49]
s: 87
i: [1,50,0,2,1,0,0,255] fc1_xboole_0
a: [3,92,0,2,1,50]
s: 92
i: [1,69,0,2,1,0,0,3] rc2_xboole_0
a: [3,102,0,2,1,69]
s: 102
i: [1,70,0,2,1,0,0,3] rc1_xboole_0
a: [3,107,0,2,1,70]
s: 107
i: [1,51,0,3,1,0,0,255] d2_xboole_0
a: [3,97,0,3,1,51]
s: 97
r: 87
i: [1,52,0,5,1,0,0,255] t2_boole
s: 52
i: [1,53,0,5,2,0,0,255] t6_boole
s: 53
i: [1,58,0,5,2,0,0,2] t7_boole
s: 58
i: [1,65,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 65
i: [1,67,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 67
i: [1,68,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 68
d: [2,113,1,5,1,1,40,68,77]
a: [3,117,1,5,1,113]
s: 117
i: [1,57,0,7,3,0,0,2] t8_boole
s: 57
i: [1,66,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 66
d: [2,120,1,5,1,0,44,66,52]
s: 120
i: [1,54,0,8,2,0,0,1] t3_xboole_0
s: 54
d: [2,121,1,5,1,0,44,66,52]
s: 121
i: [1,55,0,8,2,0,0,1] t3_xboole_0
s: 55
i: [1,72,0,8,2,0,0,2] d4_xboole_0
s: 72
d: [2,122,1,5,1,0,44,52,66]
s: 122
i: [1,73,0,8,2,0,0,2] d4_xboole_0
s: 73
d: [2,123,1,5,1,0,44,52,66]
s: 123
i: [1,56,0,9,3,0,0,1] t3_xboole_0
s: 56
i: [1,71,0,11,3,0,0,2] d4_xboole_0
s: 71
d: [2,130,1,5,2,0,40,54,58]
s: 130
i: [1,62,0,17,3,0,0,2] d4_xboole_0
s: 62
d: [2,160,2,2,1,1,40,130,82]
a: [3,165,2,2,1,160]
s: 165
i: [1,63,0,17,3,0,0,2] d4_xboole_0
s: 63
i: [1,64,0,23,4,0,0,2] d4_xboole_0
s: 64
d: [2,132,1,5,2,0,40,55,58]
s: 132
d: [2,142,1,6,2,0,44,73,52]
s: 142
d: [2,184,2,2,1,1,40,132,82]
a: [3,189,2,2,1,184]
s: 189
d: [2,110,1,7,3,0,44,53,53]
s: 110
d: [2,111,1,7,2,0,44,52,53]
s: 111
d: [2,161,2,5,2,0,40,130,68]
s: 161
d: [2,129,1,8,2,0,40,54,67]
s: 129
d: [2,139,2,6,2,0,44,72,120]
s: 139
d: [2,131,1,8,2,0,40,55,67]
s: 131
d: [2,137,1,8,2,0,44,72,66]
s: 137
d: [2,124,2,7,2,0,44,120,53]
s: 124
d: [2,138,1,8,2,0,44,72,66]
s: 138
d: [2,196,2,7,2,0,44,111,66]
s: 196
d: [2,153,1,10,3,0,40,71,58]
s: 153
d: [2,147,1,11,3,0,40,56,54]
s: 147
d: [2,230,2,9,3,0,40,147,54]
d: [2,234,2,6,2,0,18,230]
s: 234
d: [2,148,1,11,3,0,40,56,55]
s: 148
d: [2,238,2,9,3,0,40,148,55]
d: [2,240,2,6,2,0,18,238]
s: 240
d: [2,152,1,11,3,0,40,71,67]
s: 152
d: [2,172,1,11,2,0,42,62]
s: 172
d: [2,197,2,7,2,0,44,111,66]
s: 197
d: [2,178,1,11,2,0,42,63]
s: 178
d: [2,204,2,7,2,0,44,66,111]
s: 204
d: [2,133,1,12,2,0,40,72,54]
s: 133
d: [2,134,1,12,2,0,40,72,55]
s: 134
d: [2,205,2,7,2,0,44,66,111]
s: 205
d: [2,140,1,12,2,0,40,73,54]
s: 140
d: [2,275,2,7,2,0,40,133,58]
s: 275
d: [2,141,1,12,2,0,40,73,55]
s: 141
d: [2,169,1,13,3,0,40,62,58]
s: 169
d: [2,337,3,5,2,0,44,275,120]
a: [3,344,3,3,1,337]
s: 344
d: [2,175,1,13,3,0,40,63,58]
s: 175
d: [2,374,4,3,1,1,40,344,68]
s: 374
d: [2,151,1,14,4,0,40,71,56]
s: 151
d: [2,168,1,17,3,0,40,62,67]
s: 168
d: [2,395,2,9,3,1,40,151,117]
d: [2,404,2,6,2,1,18,395]
s: 404
d: [2,174,1,17,3,0,40,63,67]
s: 174
d: [2,407,3,8,2,1,40,404,54]
s: 407
d: [2,167,1,20,4,0,40,62,56]
s: 167
d: [2,173,1,20,4,0,40,63,56]
s: 173
d: [2,419,4,6,2,1,40,407,55]
d: [2,424,4,3,1,1,18,419]
a: [3,428,4,3,1,424]
s: 428
d: [2,456,5,0,0,2,52,82,428,68]
e: 456
d: [2,457,5,3,1,1,40,428,68]
d: [2,459,5,0,0,2,21,457,82]
e: 459
f: 459,424,113,47,48
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t75_xboole_1
% SZS output start Proof for xboole_1__t75_xboole_1
1. ! [X0] : ! [X1] : ~(r1_xboole_0(k3_xboole_0(X0,X1),X1) & ~r1_xboole_0(X0,X1)) [input t75_xboole_1]
2. ~! [X0] : ! [X1] : ~(r1_xboole_0(k3_xboole_0(X0,X1),X1) & ~r1_xboole_0(X0,X1)) [negated conjecture 1]
6. ! [X0] : ! [X1] : ! [X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [input d4_xboole_0]
14. ! [X0,X1] : (r1_xboole_0(X0,X1) => r1_xboole_0(X1,X0)) [input symmetry_r1_xboole_0]
16. ! [X0] : ! [X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X2] : ~(r2_hidden(X2,X1) & r2_hidden(X2,X0)) & ~r1_xboole_0(X0,X1))) [input t3_xboole_0]
20. ~! [X0,X1] : ~(r1_xboole_0(k3_xboole_0(X0,X1),X1) & ~r1_xboole_0(X0,X1)) [flattening 2]
21. ! [X0] : ! [X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X3] : ~(r2_hidden(X3,X1) & r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) [rectify 16]
22. ! [X0,X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X3] : ~(r2_hidden(X3,X1) & r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) [flattening 21]
25. ! [X0,X1,X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [flattening 6]
27. ? [X0,X1] : (r1_xboole_0(k3_xboole_0(X0,X1),X1) & ~r1_xboole_0(X0,X1)) [ennf transformation 20]
29. ! [X0,X1] : ((~r1_xboole_0(X0,X1) | ! [X2] : (~r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) & (? [X3] : (r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r1_xboole_0(X0,X1))) [ennf transformation 22]
33. ! [X0,X1] : (r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 14]
34. ? [X0,X1] : (r1_xboole_0(k3_xboole_0(X0,X1),X1) & ~r1_xboole_0(X0,X1)) => (r1_xboole_0(k3_xboole_0(sK0,sK1),sK1) & ~r1_xboole_0(sK0,sK1)) [choice axiom]
35. r1_xboole_0(k3_xboole_0(sK0,sK1),sK1) & ~r1_xboole_0(sK0,sK1) [skolemisation 27,34]
36. ! [X1,X0] : (? [X3] : (r2_hidden(X3,X1) & r2_hidden(X3,X0)) => (r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) [choice axiom]
37. ! [X0,X1] : ((~r1_xboole_0(X0,X1) | ! [X2] : (~r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) & ((r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0)) | r1_xboole_0(X0,X1))) [skolemisation 29,36]
38. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) | ~r2_hidden(X3,X0))) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [nnf transformation 25]
39. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | ~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [flattening 38]
40. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [rectify 39]
41. ! [X2,X1,X0] : (? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2))) => ((~r2_hidden(sK3(X0,X1,X2),X1) | ~r2_hidden(sK3(X0,X1,X2),X0) | ~r2_hidden(sK3(X0,X1,X2),X2)) & ((r2_hidden(sK3(X0,X1,X2),X1) & r2_hidden(sK3(X0,X1,X2),X0)) | r2_hidden(sK3(X0,X1,X2),X2)))) [choice axiom]
42. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ((~r2_hidden(sK3(X0,X1,X2),X1) | ~r2_hidden(sK3(X0,X1,X2),X0) | ~r2_hidden(sK3(X0,X1,X2),X2)) & ((r2_hidden(sK3(X0,X1,X2),X1) & r2_hidden(sK3(X0,X1,X2),X0)) | r2_hidden(sK3(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [skolemisation 40,41]
47. ~r1_xboole_0(sK0,sK1) [cnf transformation 35] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. r1_xboole_0(k3_xboole_0(sK0,sK1),sK1) [cnf transformation 35] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. r2_hidden(sK2(X0,X1),X0) | r1_xboole_0(X0,X1) [cnf transformation 37] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
55. r2_hidden(sK2(X0,X1),X1) | r1_xboole_0(X0,X1) [cnf transformation 37] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
56. ~r2_hidden(X2,X1) | ~r1_xboole_0(X0,X1) | ~r2_hidden(X2,X0) [cnf transformation 37] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
61. r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) | k3_xboole_0(X0,X1) != X2 [cnf transformation 42] {a:0,w:14,wCS:5600,sine:2,thAx:0,allAx:2,thDist:-2}
68. ~r1_xboole_0(X0,X1) | r1_xboole_0(X1,X0) [cnf transformation 33] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. r2_hidden(X4,k3_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) [equality resolution 61] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
75. 1 <=> r1_xboole_0(k3_xboole_0(sK0,sK1),sK1) [avatar definition]
77. r1_xboole_0(k3_xboole_0(sK0,sK1),sK1) <- (1) [avatar component clause 75] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
78. 1 [avatar split clause 48,75]
80. 2 <=> r1_xboole_0(sK0,sK1) [avatar definition]
82. ~r1_xboole_0(sK0,sK1) <- (~2) [avatar component clause 80] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. ~2 [avatar split clause 47,80]
113. r1_xboole_0(sK1,k3_xboole_0(sK0,sK1)) <- (1) [resolution 68,77] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
115. 8 <=> r1_xboole_0(sK1,k3_xboole_0(sK0,sK1)) [avatar definition]
117. r1_xboole_0(sK1,k3_xboole_0(sK0,sK1)) <- (8) [avatar component clause 115] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. 8 | ~1 [avatar split clause 113,75,115]
151. ~r1_xboole_0(X9,k3_xboole_0(X8,X7)) | ~r2_hidden(X6,X8) | ~r2_hidden(X6,X7) | ~r2_hidden(X6,X9) [resolution 71,56] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
395. ~r2_hidden(X16,sK0) | ~r2_hidden(X16,sK1) | ~r2_hidden(X16,sK1) <- (8) [resolution 151,117] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
404. ~r2_hidden(X16,sK1) | ~r2_hidden(X16,sK0) <- (8) [duplicate literal removal 395] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
407. ~r2_hidden(sK2(sK1,X0),sK0) | r1_xboole_0(sK1,X0) <- (8) [resolution 404,54] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
419. r1_xboole_0(sK1,sK0) | r1_xboole_0(sK1,sK0) <- (8) [resolution 407,55] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
424. r1_xboole_0(sK1,sK0) <- (8) [duplicate literal removal 419] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
426. 13 <=> r1_xboole_0(sK1,sK0) [avatar definition]
428. r1_xboole_0(sK1,sK0) <- (13) [avatar component clause 426] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
429. 13 | ~8 [avatar split clause 424,115,426]
457. r1_xboole_0(sK0,sK1) <- (13) [resolution 428,68] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
459. $false <- (~2, 13) [subsumption resolution 457,82] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
460. 2 | ~13 [avatar contradiction clause 459]
461. $false [avatar sat refutation 78,83,118,429,460] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:85,thDist:-85}
% SZS output end Proof for xboole_1__t75_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 359
% Active clauses: 60
% Passive clauses: 171
% Final active clauses: 59
% Final passive clauses: 99
% Main loop iterations started: 72
% 
% >>> Simplifying Inferences
% Duplicate literals: 43
% Fw subsumption resolutions: 5
% Fw demodulations: 10
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Forward subsumptions: 113
% 
% >>> Generating Inferences
% Binary resolution: 153
% Unit resulting resolution: 1
% Factoring: 2
% Forward superposition: 86
% Backward superposition: 23
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 7
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption: 0.001 s
% superposition: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

