i: [1,12,0,8,1,1,0,0] t9_borsuk_5
a: [3,59,0,8,1,12]
s: 59
i: [1,11,0,8,1,1,0,0] t9_borsuk_5
a: [3,64,0,8,1,11]
s: 64
i: [1,35,0,8,1,0,0,255] d9_zfmisc_1
s: 35
i: [1,36,0,8,1,0,0,255] d9_zfmisc_1
s: 36
i: [1,37,0,8,1,0,0,255] d9_zfmisc_1
s: 37
i: [1,38,0,8,1,0,0,255] d9_zfmisc_1
s: 38
i: [1,39,0,8,1,0,0,255] d9_zfmisc_1
s: 39
i: [1,40,0,8,1,0,0,255] d9_zfmisc_1
s: 40
i: [1,41,0,8,1,0,0,255] d9_zfmisc_1
s: 41
i: [1,42,0,8,1,0,0,255] d9_zfmisc_1
s: 42
i: [1,43,0,8,1,0,0,255] d9_zfmisc_1
s: 43
i: [1,44,0,8,1,0,0,255] d9_zfmisc_1
s: 44
i: [1,45,0,8,1,0,0,255] d9_zfmisc_1
s: 45
i: [1,46,0,8,1,0,0,255] d9_zfmisc_1
s: 46
i: [1,47,0,8,1,0,0,255] d9_zfmisc_1
s: 47
i: [1,48,0,8,1,0,0,255] d9_zfmisc_1
s: 48
i: [1,49,0,8,1,0,0,255] d9_zfmisc_1
s: 49
i: [1,50,0,8,1,0,0,255] d9_zfmisc_1
s: 50
i: [1,51,0,8,1,0,0,255] d9_zfmisc_1
s: 51
i: [1,52,0,8,1,0,0,255] d9_zfmisc_1
s: 52
i: [1,53,0,8,1,0,0,255] d9_zfmisc_1
s: 53
i: [1,54,0,8,1,0,0,255] d9_zfmisc_1
s: 54
i: [1,55,0,8,1,0,0,255] d9_zfmisc_1
s: 55
i: [1,34,0,71,22,0,0,255] d9_zfmisc_1
s: 34
d: [2,66,1,63,21,1,40,34,59]
a: [3,116,1,3,1,66]
s: 116
r: 64
r: 59
d: [2,194,0,8,1,2,23,64,116]
d: [2,196,0,0,0,2,21,194,52]
e: 196
r: 116
s: 59
a: [3,115,1,3,1,66]
s: 115
s: 64
a: [3,120,1,3,1,66]
s: 120
r: 64
r: 59
d: [2,200,0,8,1,2,23,64,120]
d: [2,202,0,0,0,2,21,200,54]
e: 202
r: 120
s: 59
s: 64
a: [3,119,1,3,1,66]
s: 119
a: [3,124,1,3,1,66]
s: 124
r: 64
r: 59
d: [2,207,0,8,1,2,23,64,124]
d: [2,209,0,0,0,2,21,207,51]
e: 209
r: 124
a: [3,123,1,3,1,66]
s: 123
s: 59
s: 64
a: [3,132,1,3,1,66]
s: 132
r: 64
r: 59
d: [2,216,0,8,1,2,23,64,132]
d: [2,218,0,0,0,2,21,216,53]
e: 218
r: 132
s: 59
a: [3,131,1,3,1,66]
s: 131
s: 64
a: [3,136,1,3,1,66]
s: 136
r: 64
r: 59
d: [2,225,0,8,1,2,23,64,136]
d: [2,228,0,0,0,2,21,225,47]
e: 228
r: 136
a: [3,135,1,3,1,66]
s: 135
s: 59
a: [3,144,1,3,1,66]
s: 144
r: 59
s: 64
d: [2,240,0,8,1,2,22,64,144]
d: [2,241,0,0,0,2,21,240,46]
e: 241
r: 144
s: 59
a: [3,111,1,3,1,66]
s: 111
s: 64
a: [3,140,1,3,1,66]
s: 140
r: 64
r: 59
d: [2,250,0,8,1,2,23,64,140]
d: [2,253,0,0,0,2,21,250,49]
e: 253
r: 140
s: 59
s: 64
a: [3,139,1,3,1,66]
s: 139
a: [3,143,1,3,1,66]
s: 143
a: [3,152,1,3,1,66]
s: 152
r: 64
r: 59
d: [2,263,0,8,1,2,23,64,152]
d: [2,266,0,0,0,2,21,263,48]
e: 266
r: 152
s: 59
s: 64
a: [3,151,1,3,1,66]
s: 151
a: [3,160,1,3,1,66]
s: 160
r: 64
r: 59
d: [2,278,0,8,1,2,23,64,160]
d: [2,282,0,0,0,2,21,278,37]
e: 282
r: 160
a: [3,159,1,3,1,66]
s: 159
s: 59
s: 64
a: [3,168,1,3,1,66]
s: 168
r: 64
r: 59
r: 159
d: [2,296,0,8,1,2,23,64,168]
d: [2,301,0,0,0,2,21,296,40]
e: 301
r: 168
s: 59
a: [3,148,1,3,1,66]
s: 148
r: 59
s: 64
d: [2,315,0,8,1,2,22,64,148]
d: [2,316,0,0,0,2,21,315,45]
e: 316
r: 148
s: 59
a: [3,147,1,3,1,66]
s: 147
s: 64
s: 159
a: [3,164,1,3,1,66]
s: 164
r: 64
r: 59
d: [2,329,0,8,1,2,23,64,164]
d: [2,332,0,0,0,2,21,329,36]
e: 332
r: 164
s: 59
a: [3,163,1,3,1,66]
s: 163
s: 64
a: [3,167,1,3,1,66]
s: 167
a: [3,176,1,3,1,66]
s: 176
r: 64
r: 163
r: 59
d: [2,347,0,8,1,2,23,64,176]
d: [2,351,0,0,0,2,21,347,41]
e: 351
r: 176
s: 59
s: 163
s: 64
a: [3,175,1,3,1,66]
s: 175
a: [3,184,1,3,1,66]
s: 184
r: 64
r: 59
d: [2,367,0,8,1,2,23,64,184]
d: [2,372,0,0,0,2,21,367,35]
e: 372
r: 184
s: 59
a: [3,183,1,3,1,66]
s: 183
s: 64
a: [3,192,1,3,1,66]
s: 192
r: 64
r: 183
r: 59
r: 163
d: [2,390,0,8,1,2,23,64,192]
d: [2,396,0,0,0,2,21,390,38]
e: 396
r: 192
s: 59
s: 64
a: [3,128,1,3,1,66]
s: 128
r: 64
r: 59
r: 175
r: 147
d: [2,413,0,8,1,2,23,64,128]
d: [2,417,0,0,0,2,21,413,50]
e: 417
r: 128
s: 59
a: [3,127,1,3,1,66]
s: 127
s: 64
s: 147
s: 163
a: [3,172,1,3,1,66]
s: 172
r: 64
r: 59
r: 159
d: [2,432,0,8,1,2,23,64,172]
d: [2,437,0,0,0,2,21,432,42]
e: 437
r: 172
s: 159
s: 59
s: 64
a: [3,171,1,3,1,66]
s: 171
s: 175
s: 183
a: [3,188,1,3,1,66]
s: 188
r: 59
r: 64
r: 159
r: 183
d: [2,456,0,8,1,2,23,64,188]
d: [2,462,0,0,0,2,21,456,39]
e: 462
r: 188
s: 59
a: [3,156,1,3,1,66]
s: 156
r: 59
r: 163
r: 167
s: 64
d: [2,487,0,8,1,2,22,64,156]
d: [2,488,0,0,0,2,21,487,43]
e: 488
r: 156
a: [3,155,1,3,1,66]
s: 155
s: 59
s: 64
s: 159
s: 163
s: 167
a: [3,180,1,3,1,66]
s: 180
r: 167
r: 64
r: 59
d: [2,507,0,8,1,2,23,64,180]
d: [2,511,0,0,0,2,21,507,44]
e: 511
r: 111
r: 180
s: 59
a: [3,112,1,3,1,66]
s: 112
r: 59
r: 135
r: 139
r: 143
r: 147
r: 151
s: 64
d: [2,617,0,8,1,2,22,64,112]
d: [2,618,0,0,0,2,21,617,55]
e: 618
d: [2,594,1,63,21,1,40,59,34]
d: [2,595,1,60,20,2,21,594,175]
d: [2,596,1,57,19,3,21,595,171]
d: [2,597,1,54,18,4,21,596,163]
d: [2,598,1,51,17,5,21,597,159]
d: [2,599,1,48,16,6,21,598,155]
d: [2,600,1,45,15,7,21,599,151]
d: [2,601,1,42,14,8,21,600,147]
d: [2,602,1,39,13,9,21,601,143]
d: [2,603,1,36,12,10,21,602,139]
d: [2,604,1,33,11,11,21,603,135]
d: [2,605,1,30,10,12,21,604,131]
d: [2,606,1,27,9,13,21,605,127]
d: [2,607,1,24,8,14,21,606,123]
d: [2,608,1,21,7,15,21,607,119]
d: [2,609,1,18,6,16,21,608,115]
f: 618,609,511,488,462,437,417,396,372,351,332,316,301,282,266,253,241,228,218,209,202,196,11,12
% Refutation found. Thanks to Tanya!
% SZS status Theorem for borsuk_5__t9_borsuk_5
% SZS output start Proof for borsuk_5__t9_borsuk_5
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) => r5_zfmisc_1(X0,X1,X4,X2,X5,X6,X3)) [input t9_borsuk_5]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) => r5_zfmisc_1(X0,X1,X4,X2,X5,X6,X3)) [negated conjecture 1]
3. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) <=> (X5 != X6 & X4 != X6 & X4 != X5 & X3 != X6 & X3 != X5 & X3 != X4 & X2 != X6 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X6 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X6 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1)) [input d9_zfmisc_1]
4. ~! [X0,X1,X2,X3,X4,X5,X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) => r5_zfmisc_1(X0,X1,X4,X2,X5,X6,X3)) [flattening 2]
5. ! [X0,X1,X2,X3,X4,X5,X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) <=> (X5 != X6 & X4 != X6 & X4 != X5 & X3 != X6 & X3 != X5 & X3 != X4 & X2 != X6 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X6 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X6 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1)) [flattening 3]
6. ? [X0,X1,X2,X3,X4,X5,X6] : (~r5_zfmisc_1(X0,X1,X4,X2,X5,X6,X3) & r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6)) [ennf transformation 4]
7. ? [X0,X1,X2,X3,X4,X5,X6] : (~r5_zfmisc_1(X0,X1,X4,X2,X5,X6,X3) & r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6)) => (~r5_zfmisc_1(sK0,sK1,sK4,sK2,sK5,sK6,sK3) & r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) [choice axiom]
8. ~r5_zfmisc_1(sK0,sK1,sK4,sK2,sK5,sK6,sK3) & r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6) [skolemisation 6,7]
9. ! [X0,X1,X2,X3,X4,X5,X6] : ((r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) | (X5 = X6 | X4 = X6 | X4 = X5 | X3 = X6 | X3 = X5 | X3 = X4 | X2 = X6 | X2 = X5 | X2 = X4 | X2 = X3 | X1 = X6 | X1 = X5 | X1 = X4 | X1 = X3 | X1 = X2 | X0 = X6 | X0 = X5 | X0 = X4 | X0 = X3 | X0 = X2 | X0 = X1)) & ((X5 != X6 & X4 != X6 & X4 != X5 & X3 != X6 & X3 != X5 & X3 != X4 & X2 != X6 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X6 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X6 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1) | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6))) [nnf transformation 5]
10. ! [X0,X1,X2,X3,X4,X5,X6] : ((r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) | X5 = X6 | X4 = X6 | X4 = X5 | X3 = X6 | X3 = X5 | X3 = X4 | X2 = X6 | X2 = X5 | X2 = X4 | X2 = X3 | X1 = X6 | X1 = X5 | X1 = X4 | X1 = X3 | X1 = X2 | X0 = X6 | X0 = X5 | X0 = X4 | X0 = X3 | X0 = X2 | X0 = X1) & ((X5 != X6 & X4 != X6 & X4 != X5 & X3 != X6 & X3 != X5 & X3 != X4 & X2 != X6 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X6 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X6 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1) | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6))) [flattening 9]
11. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6) [cnf transformation 8] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
12. ~r5_zfmisc_1(sK0,sK1,sK4,sK2,sK5,sK6,sK3) [cnf transformation 8] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
13. X0 != X1 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
14. X0 != X2 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
15. X0 != X3 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
16. X0 != X4 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
17. X0 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
18. X0 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
19. X1 != X2 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
20. X1 != X3 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
21. X1 != X4 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
22. X1 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
23. X1 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
24. X2 != X3 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
25. X2 != X4 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
26. X2 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
27. X2 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
28. X3 != X4 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
29. X3 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
30. X3 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
31. X4 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
32. X4 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
33. X5 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
34. r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) | X5 = X6 | X4 = X6 | X4 = X5 | X3 = X6 | X3 = X5 | X3 = X4 | X2 = X6 | X2 = X5 | X2 = X4 | X2 = X3 | X1 = X6 | X1 = X5 | X1 = X4 | X1 = X3 | X1 = X2 | X0 = X6 | X0 = X5 | X0 = X4 | X0 = X3 | X0 = X2 | X0 = X1 [cnf transformation 10] {a:0,w:71,wCS:28400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
35. ~r5_zfmisc_1(X0,X1,X2,X3,X4,X6,X6) [equality resolution 33] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
36. ~r5_zfmisc_1(X0,X1,X2,X3,X6,X5,X6) [equality resolution 32] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
37. ~r5_zfmisc_1(X0,X1,X2,X3,X5,X5,X6) [equality resolution 31] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. ~r5_zfmisc_1(X0,X1,X2,X6,X4,X5,X6) [equality resolution 30] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. ~r5_zfmisc_1(X0,X1,X2,X5,X4,X5,X6) [equality resolution 29] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. ~r5_zfmisc_1(X0,X1,X2,X4,X4,X5,X6) [equality resolution 28] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
41. ~r5_zfmisc_1(X0,X1,X6,X3,X4,X5,X6) [equality resolution 27] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. ~r5_zfmisc_1(X0,X1,X5,X3,X4,X5,X6) [equality resolution 26] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. ~r5_zfmisc_1(X0,X1,X4,X3,X4,X5,X6) [equality resolution 25] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. ~r5_zfmisc_1(X0,X1,X3,X3,X4,X5,X6) [equality resolution 24] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
45. ~r5_zfmisc_1(X0,X6,X2,X3,X4,X5,X6) [equality resolution 23] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
46. ~r5_zfmisc_1(X0,X5,X2,X3,X4,X5,X6) [equality resolution 22] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
47. ~r5_zfmisc_1(X0,X4,X2,X3,X4,X5,X6) [equality resolution 21] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
48. ~r5_zfmisc_1(X0,X3,X2,X3,X4,X5,X6) [equality resolution 20] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
49. ~r5_zfmisc_1(X0,X2,X2,X3,X4,X5,X6) [equality resolution 19] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
50. ~r5_zfmisc_1(X6,X1,X2,X3,X4,X5,X6) [equality resolution 18] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
51. ~r5_zfmisc_1(X5,X1,X2,X3,X4,X5,X6) [equality resolution 17] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
52. ~r5_zfmisc_1(X4,X1,X2,X3,X4,X5,X6) [equality resolution 16] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
53. ~r5_zfmisc_1(X3,X1,X2,X3,X4,X5,X6) [equality resolution 15] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
54. ~r5_zfmisc_1(X2,X1,X2,X3,X4,X5,X6) [equality resolution 14] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
55. ~r5_zfmisc_1(X1,X1,X2,X3,X4,X5,X6) [equality resolution 13] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
57. 1 <=> r5_zfmisc_1(sK0,sK1,sK4,sK2,sK5,sK6,sK3) [avatar definition]
59. ~r5_zfmisc_1(sK0,sK1,sK4,sK2,sK5,sK6,sK3) <- (~1) [avatar component clause 57] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. ~1 [avatar split clause 12,57]
62. 2 <=> r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6) [avatar definition]
64. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6) <- (2) [avatar component clause 62] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. 2 [avatar split clause 11,62]
110. 3 <=> sK0 = sK1 [avatar definition]
112. sK0 = sK1 <- (3) [avatar component clause 110] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 4 <=> sK0 = sK4 [avatar definition]
115. sK0 != sK4 <- (~4) [avatar component clause 114] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. sK0 = sK4 <- (4) [avatar component clause 114] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. 5 <=> sK0 = sK2 [avatar definition]
119. sK0 != sK2 <- (~5) [avatar component clause 118] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
120. sK0 = sK2 <- (5) [avatar component clause 118] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
122. 6 <=> sK0 = sK5 [avatar definition]
123. sK0 != sK5 <- (~6) [avatar component clause 122] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. sK0 = sK5 <- (6) [avatar component clause 122] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. 7 <=> sK0 = sK6 [avatar definition]
127. sK0 != sK6 <- (~7) [avatar component clause 126] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. sK0 = sK6 <- (7) [avatar component clause 126] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
130. 8 <=> sK0 = sK3 [avatar definition]
131. sK0 != sK3 <- (~8) [avatar component clause 130] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
132. sK0 = sK3 <- (8) [avatar component clause 130] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. 9 <=> sK1 = sK4 [avatar definition]
135. sK1 != sK4 <- (~9) [avatar component clause 134] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
136. sK1 = sK4 <- (9) [avatar component clause 134] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
138. 10 <=> sK1 = sK2 [avatar definition]
139. sK1 != sK2 <- (~10) [avatar component clause 138] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
140. sK1 = sK2 <- (10) [avatar component clause 138] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. 11 <=> sK1 = sK5 [avatar definition]
143. sK1 != sK5 <- (~11) [avatar component clause 142] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. sK1 = sK5 <- (11) [avatar component clause 142] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
146. 12 <=> sK1 = sK6 [avatar definition]
147. sK1 != sK6 <- (~12) [avatar component clause 146] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
148. sK1 = sK6 <- (12) [avatar component clause 146] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
150. 13 <=> sK1 = sK3 [avatar definition]
151. sK1 != sK3 <- (~13) [avatar component clause 150] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
152. sK1 = sK3 <- (13) [avatar component clause 150] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
154. 14 <=> sK2 = sK4 [avatar definition]
155. sK2 != sK4 <- (~14) [avatar component clause 154] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
156. sK2 = sK4 <- (14) [avatar component clause 154] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
158. 15 <=> sK4 = sK5 [avatar definition]
159. sK4 != sK5 <- (~15) [avatar component clause 158] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
160. sK4 = sK5 <- (15) [avatar component clause 158] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
162. 16 <=> sK4 = sK6 [avatar definition]
163. sK4 != sK6 <- (~16) [avatar component clause 162] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
164. sK4 = sK6 <- (16) [avatar component clause 162] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
166. 17 <=> sK3 = sK4 [avatar definition]
168. sK3 = sK4 <- (17) [avatar component clause 166] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
170. 18 <=> sK2 = sK5 [avatar definition]
171. sK2 != sK5 <- (~18) [avatar component clause 170] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
172. sK2 = sK5 <- (18) [avatar component clause 170] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
174. 19 <=> sK2 = sK6 [avatar definition]
175. sK2 != sK6 <- (~19) [avatar component clause 174] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
176. sK2 = sK6 <- (19) [avatar component clause 174] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
178. 20 <=> sK2 = sK3 [avatar definition]
180. sK2 = sK3 <- (20) [avatar component clause 178] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
182. 21 <=> sK5 = sK6 [avatar definition]
184. sK5 = sK6 <- (21) [avatar component clause 182] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
186. 22 <=> sK3 = sK5 [avatar definition]
188. sK3 = sK5 <- (22) [avatar component clause 186] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
190. 23 <=> sK3 = sK6 [avatar definition]
192. sK3 = sK6 <- (23) [avatar component clause 190] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
194. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK0,sK5,sK6) <- (2, 4) [backward demodulation 64,116] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
196. $false <- (2, 4) [subsumption resolution 194,52] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
197. ~2 | ~4 [avatar contradiction clause 196]
200. r5_zfmisc_1(sK0,sK1,sK0,sK3,sK4,sK5,sK6) <- (2, 5) [backward demodulation 64,120] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
202. $false <- (2, 5) [subsumption resolution 200,54] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
203. ~2 | ~5 [avatar contradiction clause 202]
207. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK0,sK6) <- (2, 6) [backward demodulation 64,124] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
209. $false <- (2, 6) [subsumption resolution 207,51] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
210. ~2 | ~6 [avatar contradiction clause 209]
216. r5_zfmisc_1(sK0,sK1,sK2,sK0,sK4,sK5,sK6) <- (2, 8) [backward demodulation 64,132] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
218. $false <- (2, 8) [subsumption resolution 216,53] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
219. ~2 | ~8 [avatar contradiction clause 218]
225. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK1,sK5,sK6) <- (2, 9) [backward demodulation 64,136] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
228. $false <- (2, 9) [subsumption resolution 225,47] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
229. ~2 | ~9 [avatar contradiction clause 228]
240. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK1,sK6) <- (2, 11) [forward demodulation 64,144] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
241. $false <- (2, 11) [subsumption resolution 240,46] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
242. ~2 | ~11 [avatar contradiction clause 241]
250. r5_zfmisc_1(sK0,sK1,sK1,sK3,sK4,sK5,sK6) <- (2, 10) [backward demodulation 64,140] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
253. $false <- (2, 10) [subsumption resolution 250,49] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
254. ~2 | ~10 [avatar contradiction clause 253]
263. r5_zfmisc_1(sK0,sK1,sK2,sK1,sK4,sK5,sK6) <- (2, 13) [backward demodulation 64,152] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
266. $false <- (2, 13) [subsumption resolution 263,48] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
267. ~2 | ~13 [avatar contradiction clause 266]
278. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK4,sK6) <- (2, 15) [backward demodulation 64,160] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
282. $false <- (2, 15) [subsumption resolution 278,37] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
283. ~2 | ~15 [avatar contradiction clause 282]
296. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK3,sK5,sK6) <- (2, 17) [backward demodulation 64,168] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
301. $false <- (2, 17) [subsumption resolution 296,40] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
302. ~2 | ~17 [avatar contradiction clause 301]
315. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK1) <- (2, 12) [forward demodulation 64,148] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
316. $false <- (2, 12) [subsumption resolution 315,45] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
317. ~2 | ~12 [avatar contradiction clause 316]
329. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK4) <- (2, 16) [backward demodulation 64,164] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
332. $false <- (2, 16) [subsumption resolution 329,36] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
333. ~2 | ~16 [avatar contradiction clause 332]
347. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK2) <- (2, 19) [backward demodulation 64,176] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
351. $false <- (2, 19) [subsumption resolution 347,41] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
352. ~2 | ~19 [avatar contradiction clause 351]
367. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK5) <- (2, 21) [backward demodulation 64,184] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
372. $false <- (2, 21) [subsumption resolution 367,35] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
373. ~2 | ~21 [avatar contradiction clause 372]
390. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK3) <- (2, 23) [backward demodulation 64,192] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
396. $false <- (2, 23) [subsumption resolution 390,38] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
397. ~2 | ~23 [avatar contradiction clause 396]
413. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK0) <- (2, 7) [backward demodulation 64,128] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
417. $false <- (2, 7) [subsumption resolution 413,50] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
418. ~2 | ~7 [avatar contradiction clause 417]
432. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK2,sK6) <- (2, 18) [backward demodulation 64,172] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
437. $false <- (2, 18) [subsumption resolution 432,42] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
438. ~2 | ~18 [avatar contradiction clause 437]
456. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK3,sK6) <- (2, 22) [backward demodulation 64,188] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
462. $false <- (2, 22) [subsumption resolution 456,39] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
463. ~2 | ~22 [avatar contradiction clause 462]
487. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK2,sK5,sK6) <- (2, 14) [forward demodulation 64,156] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
488. $false <- (2, 14) [subsumption resolution 487,43] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
489. ~2 | ~14 [avatar contradiction clause 488]
507. r5_zfmisc_1(sK0,sK1,sK2,sK2,sK4,sK5,sK6) <- (2, 20) [backward demodulation 64,180] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
511. $false <- (2, 20) [subsumption resolution 507,44] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
512. ~2 | ~20 [avatar contradiction clause 511]
594. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK2 = sK6 | sK2 = sK5 | sK3 = sK4 | sK4 = sK6 | sK4 = sK5 | sK2 = sK4 | sK1 = sK3 | sK1 = sK6 | sK1 = sK5 | sK1 = sK2 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1) [resolution 59,34] {a:1,w:63,wCS:6300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
595. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK2 = sK5 | sK3 = sK4 | sK4 = sK6 | sK4 = sK5 | sK2 = sK4 | sK1 = sK3 | sK1 = sK6 | sK1 = sK5 | sK1 = sK2 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~19) [subsumption resolution 594,175] {a:1,w:60,wCS:6000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
596. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK4 = sK6 | sK4 = sK5 | sK2 = sK4 | sK1 = sK3 | sK1 = sK6 | sK1 = sK5 | sK1 = sK2 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~18, ~19) [subsumption resolution 595,171] {a:1,w:57,wCS:5700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
597. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK4 = sK5 | sK2 = sK4 | sK1 = sK3 | sK1 = sK6 | sK1 = sK5 | sK1 = sK2 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~16, ~18, ~19) [subsumption resolution 596,163] {a:1,w:54,wCS:5400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
598. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK2 = sK4 | sK1 = sK3 | sK1 = sK6 | sK1 = sK5 | sK1 = sK2 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~15, ~16, ~18, ~19) [subsumption resolution 597,159] {a:1,w:51,wCS:5100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
599. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK1 = sK3 | sK1 = sK6 | sK1 = sK5 | sK1 = sK2 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 598,155] {a:1,w:48,wCS:4800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
600. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK1 = sK6 | sK1 = sK5 | sK1 = sK2 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 599,151] {a:1,w:45,wCS:4500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
601. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK1 = sK5 | sK1 = sK2 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~12, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 600,147] {a:1,w:42,wCS:4200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
602. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK1 = sK2 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 601,143] {a:1,w:39,wCS:3900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
603. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK1 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 602,139] {a:1,w:36,wCS:3600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
604. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK0 = sK3 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 603,135] {a:1,w:33,wCS:3300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
605. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK0 = sK6 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 604,131] {a:1,w:30,wCS:3000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
606. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK0 = sK5 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~7, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 605,127] {a:1,w:27,wCS:2700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
607. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK0 = sK2 | sK0 = sK4 | sK0 = sK1 <- (~1, ~6, ~7, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 606,123] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
608. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK0 = sK4 | sK0 = sK1 <- (~1, ~5, ~6, ~7, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 607,119] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
609. sK3 = sK6 | sK3 = sK5 | sK5 = sK6 | sK2 = sK3 | sK3 = sK4 | sK0 = sK1 <- (~1, ~4, ~5, ~6, ~7, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19) [subsumption resolution 608,115] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
610. 3 | 17 | 20 | 21 | 22 | 23 | 1 | 4 | 5 | 6 | 7 | 8 | 9 | 10 | 11 | 12 | 13 | 14 | 15 | 16 | 18 | 19 [avatar split clause 609,174,170,162,158,154,150,146,142,138,134,130,126,122,118,114,57,190,186,182,178,166,110]
617. r5_zfmisc_1(sK0,sK0,sK2,sK3,sK4,sK5,sK6) <- (2, 3) [forward demodulation 64,112] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
618. $false <- (2, 3) [subsumption resolution 617,55] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
619. ~2 | ~3 [avatar contradiction clause 618]
620. $false [avatar sat refutation 60,65,197,203,210,219,229,242,254,267,283,302,317,333,352,373,397,418,438,463,489,512,610,619] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:913,thDist:-913}
% SZS output end Proof for borsuk_5__t9_borsuk_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 2
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 547
% Active clauses: 109
% Passive clauses: 119
% Final active clauses: 33
% Final passive clauses: 5
% Main loop iterations started: 113
% 
% >>> Simplifying Inferences
% Duplicate literals: 105
% Fw subsumption resolutions: 267
% Fw demodulations: 4
% Bw demodulations: 55
% 
% >>> Deletion Inferences
% Equational tautologies: 21
% Forward subsumptions: 51
% 
% >>> Generating Inferences
% Binary resolution: 42
% Backward superposition: 24
% 
% >>> AVATAR
% Split clauses: 34
% Split components: 335
% Sat splitting refutations: 21
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 61
% SAT solver unit clauses: 2
% SAT solver binary clauses: 21
% 
% Memory used [KB]: 10106
% Time elapsed: 0.647 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.137 s
% congruence closure: 0.001 s
% forward subsumption: 0.002 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.640 s ( own 0.503 s ) 
% other: 0.647 s ( own 0.003 s ) 

