i: [1,46,0,4,1,1,0,0] t76_zfmisc_1
a: [3,72,0,4,1,46]
s: 72
i: [1,45,0,6,2,1,0,0] t76_zfmisc_1
s: 45
i: [1,47,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,77,0,2,1,47]
s: 77
i: [1,48,0,2,1,0,0,255] fc1_xboole_0
a: [3,82,0,2,1,48]
s: 82
i: [1,64,0,2,1,0,0,2] rc2_xboole_0
a: [3,92,0,2,1,64]
s: 92
i: [1,65,0,2,1,0,0,2] rc1_xboole_0
a: [3,97,0,2,1,65]
s: 97
i: [1,62,0,3,1,0,0,1] reflexivity_r1_tarski
s: 62
i: [1,49,0,3,1,0,0,255] d2_xboole_0
a: [3,87,0,3,1,49]
s: 87
r: 77
i: [1,50,0,5,2,0,0,255] t6_boole
s: 50
i: [1,61,0,5,2,0,0,1] t7_boole
s: 61
i: [1,63,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 63
i: [1,60,0,7,3,0,0,2] t8_boole
s: 60
i: [1,52,0,8,2,0,0,1] d3_tarski
s: 52
d: [2,102,1,8,2,0,40,52,45]
s: 102
i: [1,53,0,8,2,0,0,1] d3_tarski
s: 53
i: [1,51,0,9,3,0,0,1] d3_tarski
s: 51
d: [2,108,2,11,3,0,40,51,102]
s: 108
i: [1,67,0,9,2,0,0,1] d4_tarski
s: 67
d: [2,110,1,9,2,0,40,67,45]
s: 110
i: [1,68,0,9,2,0,0,1] d4_tarski
s: 68
i: [1,66,0,10,3,0,0,1] d4_tarski
s: 66
d: [2,114,2,12,3,0,40,110,51]
s: 114
i: [1,58,0,14,3,0,0,1] d4_tarski
s: 58
d: [2,120,3,11,3,0,40,114,68]
d: [2,123,3,7,2,0,18,120]
s: 123
i: [1,57,0,16,3,0,0,1] d4_tarski
s: 57
i: [1,59,0,17,4,0,0,1] d4_tarski
s: 59
d: [2,129,4,9,3,0,40,123,66]
s: 129
d: [2,126,1,14,3,0,40,58,45]
s: 126
d: [2,130,4,10,2,0,40,123,52]
s: 130
d: [2,143,5,8,2,0,40,130,53]
d: [2,146,5,4,1,0,18,143]
d: [2,147,5,0,0,1,21,146,72]
e: 147
f: 147,46
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t76_zfmisc_1
% SZS output start Proof for zfmisc_1__t76_zfmisc_1
1. ! [X0] : ! [X1] : (! [X2] : (r2_hidden(X2,X0) => r1_tarski(X2,X1)) => r1_tarski(k3_tarski(X0),X1)) [input t76_zfmisc_1]
2. ~! [X0] : ! [X1] : (! [X2] : (r2_hidden(X2,X0) => r1_tarski(X2,X1)) => r1_tarski(k3_tarski(X0),X1)) [negated conjecture 1]
5. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
6. ! [X0] : ! [X1] : (k3_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)))) [input d4_tarski]
17. ~! [X0,X1] : (! [X2] : (r2_hidden(X2,X0) => r1_tarski(X2,X1)) => r1_tarski(k3_tarski(X0),X1)) [flattening 2]
18. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 5]
19. ! [X0,X1] : (k3_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)))) [flattening 6]
23. ? [X0,X1] : (~r1_tarski(k3_tarski(X0),X1) & ! [X2] : (r1_tarski(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 17]
25. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 18]
29. ? [X0,X1] : (~r1_tarski(k3_tarski(X0),X1) & ! [X2] : (r1_tarski(X2,X1) | ~r2_hidden(X2,X0))) => (~r1_tarski(k3_tarski(sK0),sK1) & ! [X2] : (r1_tarski(X2,sK1) | ~r2_hidden(X2,sK0))) [choice axiom]
30. ~r1_tarski(k3_tarski(sK0),sK1) & ! [X2] : (r1_tarski(X2,sK1) | ~r2_hidden(X2,sK0)) [skolemisation 23,29]
31. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 25]
32. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 31]
33. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) [choice axiom]
34. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 32,33]
35. ! [X0,X1] : ((k3_tarski(X0) = X1 | ? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3))) & (? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)) | ~r2_hidden(X2,X1))) | k3_tarski(X0) != X1)) [nnf transformation 19]
36. ! [X0,X1] : ((k3_tarski(X0) = X1 | ? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(X2,X4)) | r2_hidden(X2,X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (~r2_hidden(X6,X0) | ~r2_hidden(X5,X6))) & (? [X7] : (r2_hidden(X7,X0) & r2_hidden(X5,X7)) | ~r2_hidden(X5,X1))) | k3_tarski(X0) != X1)) [rectify 35]
37. ! [X1,X0] : (? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(X2,X4)) | r2_hidden(X2,X1))) => ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(sK3(X0,X1),X3)) | ~r2_hidden(sK3(X0,X1),X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(sK3(X0,X1),X4)) | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
38. ! [X1,X0] : (? [X4] : (r2_hidden(X4,X0) & r2_hidden(sK3(X0,X1),X4)) => (r2_hidden(sK4(X0,X1),X0) & r2_hidden(sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
39. ! [X5,X0] : (? [X7] : (r2_hidden(X7,X0) & r2_hidden(X5,X7)) => (r2_hidden(sK5(X0,X5),X0) & r2_hidden(X5,sK5(X0,X5)))) [choice axiom]
40. ! [X0,X1] : ((k3_tarski(X0) = X1 | ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(sK3(X0,X1),X3)) | ~r2_hidden(sK3(X0,X1),X1)) & ((r2_hidden(sK4(X0,X1),X0) & r2_hidden(sK3(X0,X1),sK4(X0,X1))) | r2_hidden(sK3(X0,X1),X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (~r2_hidden(X6,X0) | ~r2_hidden(X5,X6))) & ((r2_hidden(sK5(X0,X5),X0) & r2_hidden(X5,sK5(X0,X5))) | ~r2_hidden(X5,X1))) | k3_tarski(X0) != X1)) [skolemisation 36,39,38,37]
45. ~r2_hidden(X2,sK0) | r1_tarski(X2,sK1) [cnf transformation 30] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. ~r1_tarski(k3_tarski(sK0),sK1) [cnf transformation 30] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. ~r1_tarski(X0,X1) | ~r2_hidden(X3,X0) | r2_hidden(X3,X1) [cnf transformation 34] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
52. r2_hidden(sK2(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 34] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
53. ~r2_hidden(sK2(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 34] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
54. r2_hidden(X5,sK5(X0,X5)) | ~r2_hidden(X5,X1) | k3_tarski(X0) != X1 [cnf transformation 40] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:4,thDist:-4}
55. r2_hidden(sK5(X0,X5),X0) | ~r2_hidden(X5,X1) | k3_tarski(X0) != X1 [cnf transformation 40] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:4,thDist:-4}
67. r2_hidden(sK5(X0,X5),X0) | ~r2_hidden(X5,k3_tarski(X0)) [equality resolution 55] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
68. r2_hidden(X5,sK5(X0,X5)) | ~r2_hidden(X5,k3_tarski(X0)) [equality resolution 54] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
70. 1 <=> r1_tarski(k3_tarski(sK0),sK1) [avatar definition]
72. ~r1_tarski(k3_tarski(sK0),sK1) <- (~1) [avatar component clause 70] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
73. ~1 [avatar split clause 46,70]
110. r1_tarski(sK5(sK0,X0),sK1) | ~r2_hidden(X0,k3_tarski(sK0)) [resolution 67,45] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
114. ~r2_hidden(X1,sK5(sK0,X0)) | ~r2_hidden(X0,k3_tarski(sK0)) | r2_hidden(X1,sK1) [resolution 110,51] {a:2,w:12,wCS:1200,nSel:1,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
120. ~r2_hidden(X0,k3_tarski(sK0)) | r2_hidden(X0,sK1) | ~r2_hidden(X0,k3_tarski(sK0)) [resolution 114,68] {a:3,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:12,thDist:-12}
123. ~r2_hidden(X0,k3_tarski(sK0)) | r2_hidden(X0,sK1) [duplicate literal removal 120] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:0,thAx:0,allAx:12,thDist:-12}
130. r2_hidden(sK2(k3_tarski(sK0),X2),sK1) | r1_tarski(k3_tarski(sK0),X2) [resolution 123,52] {a:4,w:10,wCS:1000,nSel:1,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
143. r1_tarski(k3_tarski(sK0),sK1) | r1_tarski(k3_tarski(sK0),sK1) [resolution 130,53] {a:5,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
146. r1_tarski(k3_tarski(sK0),sK1) [duplicate literal removal 143] {a:5,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
147. $false <- (~1) [subsumption resolution 146,72] {a:5,w:0,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
148. 1 [avatar contradiction clause 147]
149. $false [avatar sat refutation 73,148] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:31,thDist:-31}
% SZS output end Proof for zfmisc_1__t76_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 76
% Active clauses: 29
% Passive clauses: 55
% Final active clauses: 28
% Final passive clauses: 26
% Main loop iterations started: 29
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 41
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

