i: [1,38,0,3,1,1,0,0] t5_arytm_1
a: [3,51,0,3,1,38]
s: 51
i: [1,37,0,3,1,1,0,0] t5_arytm_1
a: [3,56,0,3,1,37]
s: 56
i: [1,36,0,3,1,1,0,0] t5_arytm_1
a: [3,61,0,3,1,36]
s: 61
i: [1,35,0,3,1,1,0,0] t5_arytm_1
a: [3,66,0,3,1,35]
s: 66
i: [1,34,0,3,1,1,0,0] t5_arytm_1
a: [3,71,0,3,1,34]
s: 71
d: [2,83,0,3,1,2,22,61,56]
a: [3,87,0,3,1,83]
s: 87
d: [2,89,0,3,1,2,22,66,56]
a: [3,93,0,3,1,89]
s: 93
i: [1,39,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,76,0,3,1,39]
s: 76
i: [1,40,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,81,0,3,1,40]
s: 81
i: [1,44,0,4,1,0,0,1] existence_m1_subset_1
s: 44
i: [1,45,0,11,3,0,0,2] dt_k7_arytm_2
s: 45
i: [1,47,0,12,4,0,0,1] connectedness_r1_arytm_2
s: 47
i: [1,46,0,13,3,0,0,2] commutativity_k7_arytm_2
s: 46
d: [2,95,1,9,3,1,40,47,71]
s: 95
i: [1,41,0,14,4,0,0,1] t5_arytm_2
s: 41
i: [1,42,0,14,4,0,0,1] t9_arytm_2
s: 42
d: [2,103,2,6,2,1,40,95,71]
d: [2,107,2,3,1,1,18,103]
a: [3,111,2,3,1,107]
s: 111
i: [1,43,0,16,4,0,0,1] t9_arytm_2
s: 43
d: [2,106,2,8,2,1,40,95,44]
a: [3,120,2,4,1,106]
s: 120
d: [2,96,1,9,3,1,40,47,93]
s: 96
d: [2,98,1,11,3,0,40,47,44]
s: 98
d: [2,133,2,6,2,1,40,96,93]
d: [2,138,2,3,1,1,18,133]
a: [3,142,2,3,1,138]
s: 142
d: [2,123,1,20,6,0,44,41,46]
d: [2,124,1,14,4,0,18,123]
s: 124
d: [2,137,2,8,2,1,40,96,44]
a: [3,147,2,4,1,137]
s: 147
d: [2,122,1,20,6,0,44,41,46]
d: [2,125,1,14,4,0,18,122]
s: 125
d: [2,128,1,23,7,0,44,41,43]
d: [2,131,1,20,6,0,18,128]
d: [2,132,1,15,5,0,21,131,42]
s: 132
d: [2,167,2,12,4,0,47,132]
d: [2,168,2,9,3,1,21,167,93]
s: 168
d: [2,169,3,0,0,4,52,51,71,87,168]
e: 169
d: [2,170,3,6,2,2,40,168,87]
d: [2,175,3,3,1,3,21,170,71]
d: [2,176,3,0,0,4,21,175,51]
e: 176
f: 176,89,83,34,35,36,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t5_arytm_1
% SZS output start Proof for arytm_1__t5_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((k11_arytm_3 = X1 & r1_arytm_2(X0,X1)) => k11_arytm_3 = X0))) [input t5_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((k11_arytm_3 = X1 & r1_arytm_2(X0,X1)) => k11_arytm_3 = X0))) [negated conjecture 1]
13. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = k7_arytm_2(X0,X1) => k11_arytm_3 = X0))) [input t5_arytm_2]
14. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ~(! [X2] : (m1_subset_1(X2,k2_arytm_2) => k7_arytm_2(X0,X2) != X1) & r1_arytm_2(X0,X1)))) [input t9_arytm_2]
15. ? [X0] : (? [X1] : ((k11_arytm_3 != X0 & (k11_arytm_3 = X1 & r1_arytm_2(X0,X1))) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
16. ? [X0] : (? [X1] : (k11_arytm_3 != X0 & k11_arytm_3 = X1 & r1_arytm_2(X0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 15]
17. ! [X0] : (! [X1] : ((k11_arytm_3 = X0 | k11_arytm_3 != k7_arytm_2(X0,X1)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 13]
18. ! [X0] : (! [X1] : (k11_arytm_3 = X0 | k11_arytm_3 != k7_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 17]
19. ! [X0] : (! [X1] : ((? [X2] : (k7_arytm_2(X0,X2) = X1 & m1_subset_1(X2,k2_arytm_2)) | ~r1_arytm_2(X0,X1)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 14]
20. ! [X0] : (! [X1] : (? [X2] : (k7_arytm_2(X0,X2) = X1 & m1_subset_1(X2,k2_arytm_2)) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 19]
27. ? [X0] : (? [X1] : (k11_arytm_3 != X0 & k11_arytm_3 = X1 & r1_arytm_2(X0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (k11_arytm_3 != sK0 & k11_arytm_3 = X1 & r1_arytm_2(sK0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
28. ? [X1] : (k11_arytm_3 != sK0 & k11_arytm_3 = X1 & r1_arytm_2(sK0,X1) & m1_subset_1(X1,k2_arytm_2)) => (k11_arytm_3 != sK0 & k11_arytm_3 = sK1 & r1_arytm_2(sK0,sK1) & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
29. (k11_arytm_3 != sK0 & k11_arytm_3 = sK1 & r1_arytm_2(sK0,sK1) & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 16,28,27]
30. ! [X1,X0] : (? [X2] : (k7_arytm_2(X0,X2) = X1 & m1_subset_1(X2,k2_arytm_2)) => (k7_arytm_2(X0,sK2(X0,X1)) = X1 & m1_subset_1(sK2(X0,X1),k2_arytm_2))) [choice axiom]
31. ! [X0] : (! [X1] : ((k7_arytm_2(X0,sK2(X0,X1)) = X1 & m1_subset_1(sK2(X0,X1),k2_arytm_2)) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [skolemisation 20,30]
34. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 29] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
35. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 29] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
36. r1_arytm_2(sK0,sK1) [cnf transformation 29] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
37. k11_arytm_3 = sK1 [cnf transformation 29] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. k11_arytm_3 != sK0 [cnf transformation 29] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. k11_arytm_3 != k7_arytm_2(X0,X1) | k11_arytm_3 = X0 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 18] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. m1_subset_1(sK2(X0,X1),k2_arytm_2) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 31] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
43. k7_arytm_2(X0,sK2(X0,X1)) = X1 | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 31] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
49. 1 <=> k11_arytm_3 = sK0 [avatar definition]
51. k11_arytm_3 != sK0 <- (~1) [avatar component clause 49] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
52. ~1 [avatar split clause 38,49]
54. 2 <=> k11_arytm_3 = sK1 [avatar definition]
56. k11_arytm_3 = sK1 <- (2) [avatar component clause 54] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
57. 2 [avatar split clause 37,54]
59. 3 <=> r1_arytm_2(sK0,sK1) [avatar definition]
61. r1_arytm_2(sK0,sK1) <- (3) [avatar component clause 59] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
62. 3 [avatar split clause 36,59]
64. 4 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
66. m1_subset_1(sK1,k2_arytm_2) <- (4) [avatar component clause 64] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
67. 4 [avatar split clause 35,64]
69. 5 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
71. m1_subset_1(sK0,k2_arytm_2) <- (5) [avatar component clause 69] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. 5 [avatar split clause 34,69]
83. r1_arytm_2(sK0,k11_arytm_3) <- (2, 3) [forward demodulation 61,56] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
85. 8 <=> r1_arytm_2(sK0,k11_arytm_3) [avatar definition]
87. r1_arytm_2(sK0,k11_arytm_3) <- (8) [avatar component clause 85] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
88. 8 | ~2 | ~3 [avatar split clause 83,59,54,85]
89. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (2, 4) [forward demodulation 66,56] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. 9 <=> m1_subset_1(k11_arytm_3,k2_arytm_2) [avatar definition]
93. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (9) [avatar component clause 91] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
94. 9 | ~2 | ~4 [avatar split clause 89,64,54,91]
128. k11_arytm_3 != X1 | k11_arytm_3 = X0 | ~m1_subset_1(sK2(X0,X1),k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [superposition 41,43] {a:1,w:23,wCS:9200,sine:1,thAx:0,allAx:3,thDist:-3}
131. k11_arytm_3 != X1 | k11_arytm_3 = X0 | ~m1_subset_1(sK2(X0,X1),k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) [duplicate literal removal 128] {a:1,w:20,wCS:8000,sine:1,thAx:0,allAx:3,thDist:-3}
132. k11_arytm_3 != X1 | k11_arytm_3 = X0 | ~m1_subset_1(X0,k2_arytm_2) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) [subsumption resolution 131,42] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
167. k11_arytm_3 = X0 | ~m1_subset_1(X0,k2_arytm_2) | ~r1_arytm_2(X0,k11_arytm_3) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) [equality resolution 132] {a:2,w:12,wCS:4800,sine:1,thAx:0,allAx:3,thDist:-3}
168. ~r1_arytm_2(X0,k11_arytm_3) | ~m1_subset_1(X0,k2_arytm_2) | k11_arytm_3 = X0 <- (9) [subsumption resolution 167,93] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
170. ~m1_subset_1(sK0,k2_arytm_2) | k11_arytm_3 = sK0 <- (8, 9) [resolution 168,87] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
175. k11_arytm_3 = sK0 <- (5, 8, 9) [subsumption resolution 170,71] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
176. $false <- (~1, 5, 8, 9) [subsumption resolution 175,51] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
177. 1 | ~5 | ~8 | ~9 [avatar contradiction clause 176]
181. $false [avatar sat refutation 52,57,62,67,72,88,94,177] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:94,thDist:-94}
% SZS output end Proof for arytm_1__t5_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 14
% Generated clauses: 81
% Active clauses: 24
% Passive clauses: 35
% Final active clauses: 24
% Final passive clauses: 8
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Duplicate literals: 17
% Fw subsumption resolutions: 9
% Fw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Equational tautologies: 1
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 19
% Unit resulting resolution: 1
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 4
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 8
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

