i: [1,182,0,7,2,1,0,0] t3_waybel_7
a: [3,255,0,2,1,182]
s: 255
i: [1,180,0,2,1,1,0,0] t3_waybel_7
a: [3,265,0,2,1,180]
s: 265
i: [1,179,0,2,1,1,0,0] t3_waybel_7
a: [3,270,0,2,1,179]
s: 270
i: [1,178,0,2,1,1,0,0] t3_waybel_7
a: [3,275,0,2,1,178]
s: 275
i: [1,177,0,2,1,1,0,0] t3_waybel_7
a: [3,280,0,2,1,177]
s: 280
a: [3,258,0,5,1,182]
s: 258
i: [1,183,0,2,1,0,0,255] spc1_boole
a: [3,285,0,2,1,183]
s: 285
i: [1,184,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,290,0,2,1,184]
s: 290
i: [1,185,0,2,1,0,0,255] fc1_xboole_0
a: [3,295,0,2,1,185]
s: 295
i: [1,240,0,2,1,0,0,2] rc2_xboole_0
a: [3,305,0,2,1,240]
s: 305
i: [1,241,0,2,1,0,0,1] existence_l1_orders_2
a: [3,310,0,2,1,241]
s: 310
i: [1,242,0,2,1,0,0,2] existence_l1_struct_0
a: [3,315,0,2,1,242]
s: 315
i: [1,243,0,2,1,0,0,2] rc1_xboole_0
a: [3,320,0,2,1,243]
s: 320
i: [1,245,0,2,1,0,0,255] rc1_relat_1
a: [3,325,0,2,1,245]
s: 325
i: [1,244,0,2,1,0,0,255] rc1_relat_1
a: [3,330,0,2,1,244]
s: 330
i: [1,247,0,2,1,0,0,255] rc1_finset_1
a: [3,335,0,2,1,247]
s: 335
i: [1,246,0,2,1,0,0,255] rc1_finset_1
a: [3,340,0,2,1,246]
s: 340
i: [1,250,0,2,1,0,0,255] rc8_finset_1
a: [3,345,0,2,1,250]
s: 345
i: [1,249,0,2,1,0,0,255] rc8_finset_1
a: [3,350,0,2,1,249]
s: 350
i: [1,248,0,2,1,0,0,255] rc8_finset_1
a: [3,355,0,2,1,248]
s: 355
i: [1,251,0,2,1,0,0,255] rc2_relat_1
a: [3,360,0,2,1,251]
s: 360
i: [1,186,0,3,1,0,0,255] d2_xboole_0
a: [3,300,0,3,1,186]
s: 300
r: 290
i: [1,187,0,4,2,0,0,255] cc5_finset_1
s: 187
i: [1,188,0,4,2,0,0,255] cc4_finset_1
s: 188
i: [1,190,0,4,2,0,0,1] dt_l1_orders_2
s: 190
i: [1,216,0,4,2,0,0,255] cc6_finset_1
s: 216
d: [2,363,1,2,1,1,40,190,265]
a: [3,368,1,2,1,363]
s: 368
i: [1,217,0,4,2,0,0,255] cc1_relat_1
s: 217
d: [2,364,1,2,1,1,40,190,310]
a: [3,373,1,2,1,364]
s: 373
i: [1,218,0,4,2,0,0,255] cc1_finset_1
s: 218
i: [1,236,0,4,1,0,0,2] existence_m1_subset_1
s: 236
d: [2,375,1,2,1,1,40,216,320]
a: [3,380,1,2,1,375]
s: 380
i: [1,219,0,5,2,0,0,255] t6_boole
s: 219
d: [2,376,1,2,1,1,40,216,295]
a: [3,385,1,2,1,376]
s: 385
i: [1,235,0,5,2,0,0,255] t7_boole
s: 235
i: [1,195,0,6,3,0,0,1] cc4_yellow_0
s: 195
d: [2,414,1,4,2,1,40,195,265]
d: [2,416,1,2,1,2,21,414,270]
a: [3,420,1,2,1,416]
s: 420
i: [1,196,0,6,3,0,0,1] cc4_yellow_0
s: 196
d: [2,431,1,4,2,1,40,196,265]
d: [2,433,1,2,1,2,21,431,270]
a: [3,437,1,2,1,433]
s: 437
i: [1,198,0,6,3,0,0,255] cc7_struct_0
s: 198
i: [1,199,0,6,3,0,0,255] cc5_struct_0
s: 199
d: [2,387,1,2,1,1,40,217,320]
a: [3,392,1,2,1,387]
s: 392
i: [1,201,0,6,3,0,0,1] cc2_struct_0
s: 201
d: [2,388,1,2,1,1,40,217,295]
a: [3,397,1,2,1,388]
s: 397
i: [1,202,0,6,3,0,0,1] cc17_struct_0
s: 202
i: [1,203,0,6,3,0,0,255] cc6_struct_0
s: 203
d: [2,399,1,2,1,1,40,218,320]
a: [3,404,1,2,1,399]
s: 404
i: [1,204,0,6,3,0,0,1] cc1_struct_0
s: 204
d: [2,400,1,2,1,1,40,218,295]
a: [3,409,1,2,1,400]
s: 409
i: [1,207,0,6,3,0,0,255] cc4_struct_0
s: 207
i: [1,232,0,6,2,0,0,255] t1_subset
s: 232
d: [2,415,1,4,2,1,40,195,310]
a: [3,429,1,2,1,415]
s: 429
i: [1,239,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 239
d: [2,444,1,4,2,1,40,198,315]
a: [3,450,1,2,1,444]
s: 450
i: [1,189,0,7,3,0,0,255] cc7_finset_1
s: 189
i: [1,191,0,7,2,0,0,1] dt_k4_yellow_0
s: 191
d: [2,465,1,4,2,1,40,199,315]
a: [3,471,1,2,1,465]
s: 471
i: [1,192,0,7,2,0,0,1] dt_k3_yellow_0
s: 192
d: [2,446,2,4,2,1,40,198,373]
a: [3,459,2,2,1,446]
s: 459
i: [1,205,0,7,3,0,0,255] cc8_struct_0
s: 205
i: [1,208,0,7,3,0,0,255] cc11_struct_0
s: 208
d: [2,467,2,4,2,1,40,199,373]
a: [3,476,2,2,1,467]
s: 476
i: [1,209,0,7,3,0,0,255] cc11_struct_0
s: 209
d: [2,485,1,5,2,0,40,189,236]
s: 485
i: [1,210,0,7,3,0,0,255] cc9_struct_0
s: 210
i: [1,220,0,7,3,0,0,255] fc9_struct_0
s: 220
i: [1,221,0,7,3,0,0,255] fc6_struct_0
s: 221
i: [1,224,0,7,3,0,0,2] fc2_struct_0
s: 224
i: [1,226,0,7,3,0,0,255] fc8_struct_0
s: 226
i: [1,227,0,7,3,0,0,255] fc7_struct_0
s: 227
i: [1,231,0,7,3,0,0,2] fc1_struct_0
s: 231
i: [1,234,0,7,3,0,0,1] t8_boole
s: 234
i: [1,193,0,8,2,0,0,255] d11_yellow_0
s: 193
d: [2,411,1,7,3,0,44,219,219]
s: 411
i: [1,194,0,8,2,0,0,255] d12_yellow_0
s: 194
i: [1,197,0,8,4,0,0,1] cc5_yellow_0
s: 197
d: [2,496,1,7,3,0,40,227,187]
s: 496
i: [1,233,0,8,3,0,0,255] t2_subset
s: 233
d: [2,508,1,6,2,0,40,233,236]
s: 508
i: [1,237,0,8,2,0,0,255] dt_k1_yellow_0
s: 237
i: [1,238,0,8,2,0,0,255] dt_k2_yellow_0
s: 238
d: [2,511,2,6,2,0,40,508,239]
s: 511
i: [1,211,0,9,4,0,0,255] cc10_struct_0
s: 211
d: [2,499,1,7,3,0,40,231,218]
s: 499
i: [1,213,0,9,3,0,0,1] d10_struct_0
s: 213
i: [1,214,0,9,3,0,0,1] d10_struct_0
s: 214
d: [2,500,1,7,3,0,40,231,217]
s: 500
i: [1,215,0,9,3,0,0,1] d10_struct_0
s: 215
d: [2,501,1,7,3,0,40,231,216]
s: 501
i: [1,212,0,15,5,0,0,1] d10_struct_0
s: 212
i: [1,222,0,17,6,0,0,2] t45_yellow_0
s: 222
d: [2,486,1,8,3,0,40,191,189]
s: 486
i: [1,223,0,17,6,0,0,2] t44_yellow_0
s: 223
d: [2,547,1,17,6,1,44,223,258]
d: [2,548,1,15,5,2,21,547,265]
d: [2,549,1,13,4,3,21,548,420]
d: [2,550,1,11,3,4,21,549,275]
d: [2,551,1,9,2,5,21,550,280]
s: 551
i: [1,225,0,23,7,0,0,3] t2_orders_2
s: 225
d: [2,487,1,8,3,0,40,192,189]
s: 487
d: [2,554,2,26,7,5,40,225,551]
d: [2,555,2,22,6,5,18,554]
d: [2,560,2,20,5,5,21,555,265]
d: [2,561,2,18,4,5,21,560,275]
a: [3,565,2,5,1,561]
s: 565
d: [2,572,3,0,0,2,52,265,565,191]
e: 572
d: [2,573,3,2,1,1,40,565,191]
d: [2,575,3,0,0,2,21,573,265]
e: 575
r: 565
d: [2,489,1,9,4,0,44,205,219]
s: 489
a: [3,564,2,5,1,561]
s: 564
d: [2,491,1,9,4,0,44,210,219]
s: 491
d: [2,515,1,9,3,0,40,237,189]
s: 515
d: [2,585,3,6,2,1,40,564,189]
a: [3,598,3,3,1,585]
s: 598
d: [2,519,1,9,3,0,40,238,189]
s: 519
d: [2,584,3,8,2,1,40,564,233]
a: [3,589,3,5,1,584]
s: 589
d: [2,503,1,10,3,0,44,193,219]
s: 503
d: [2,504,1,10,3,0,44,194,219]
s: 504
a: [3,592,3,3,1,584]
s: 592
d: [2,509,1,10,3,0,40,233,191]
s: 509
d: [2,615,4,5,1,1,40,589,239]
a: [3,621,4,5,1,615]
s: 621
d: [2,510,1,10,3,0,40,233,192]
s: 510
d: [2,529,1,10,4,0,40,213,189]
s: 529
a: [3,568,2,13,3,561]
s: 568
d: [2,531,1,10,4,0,40,214,189]
s: 531
d: [2,643,3,20,7,1,40,568,222]
d: [2,645,3,16,6,1,18,643]
d: [2,647,3,14,5,2,21,645,265]
d: [2,648,3,12,4,3,21,647,437]
d: [2,649,3,10,3,4,21,648,275]
d: [2,650,3,8,2,5,21,649,280]
s: 650
r: 568
r: 255
r: 551
r: 258
r: 650
s: 255
s: 258
d: [2,658,4,9,3,5,40,650,213]
d: [2,663,4,7,2,6,21,658,255]
d: [2,664,4,5,1,7,21,663,368]
a: [3,668,4,5,1,664]
s: 668
d: [2,693,1,17,6,1,44,223,258]
d: [2,695,1,15,5,2,21,693,265]
d: [2,696,1,13,4,3,21,695,420]
d: [2,697,1,11,3,4,21,696,275]
d: [2,698,1,9,2,5,21,697,280]
s: 698
d: [2,659,4,9,3,5,40,650,214]
d: [2,670,4,7,2,6,21,659,255]
d: [2,671,4,5,1,7,21,670,368]
a: [3,675,4,5,1,671]
s: 675
d: [2,706,5,9,3,1,44,215,675]
d: [2,708,5,4,2,2,21,706,668]
d: [2,709,5,2,1,3,21,708,255]
d: [2,710,5,0,0,4,21,709,368]
e: 710
r: 255
r: 698
r: 258
a: [3,254,0,2,1,182]
s: 254
a: [3,259,0,5,1,182]
s: 259
d: [2,445,2,4,2,1,40,198,368]
a: [3,690,2,2,1,445]
s: 690
d: [2,514,1,11,3,0,40,237,233]
s: 514
d: [2,660,4,6,1,5,40,650,236]
a: [3,680,4,6,1,660]
s: 680
d: [2,518,1,11,3,0,40,238,233]
s: 518
d: [2,528,1,12,4,0,40,213,233]
s: 528
d: [2,583,3,12,4,1,40,564,212]
d: [2,682,3,10,3,2,21,583,368]
a: [3,685,3,8,2,682]
s: 685
d: [2,738,4,7,2,1,40,685,192]
d: [2,744,4,2,1,2,21,738,259]
d: [2,745,4,0,0,3,21,744,265]
e: 745
i: [1,181,0,7,2,1,0,0] t3_waybel_7
d: [2,700,2,26,7,5,40,698,225]
d: [2,701,2,22,6,5,18,700]
d: [2,702,2,17,5,6,21,701,564]
d: [2,703,2,15,4,6,21,702,265]
d: [2,704,2,13,3,6,21,703,275]
d: [2,712,5,2,1,3,21,708,368]
f: 745,712,704,682,671,664,575,433,416,363,177,178,179,180,181,182
% Refutation found. Thanks to Tanya!
% SZS status Theorem for waybel_7__t3_waybel_7
% SZS output start Proof for waybel_7__t3_waybel_7
1. ! [X0] : ((l1_orders_2(X0) & v3_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) => (v7_struct_0(X0) <=> k4_yellow_0(X0) = k3_yellow_0(X0))) [input t3_waybel_7]
2. ~! [X0] : ((l1_orders_2(X0) & v3_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) => (v7_struct_0(X0) <=> k4_yellow_0(X0) = k3_yellow_0(X0))) [negated conjecture 1]
17. ! [X0] : (l1_orders_2(X0) => (v3_yellow_0(X0) => (v2_yellow_0(X0) & v1_yellow_0(X0)))) [input cc4_yellow_0]
28. ! [X0] : (l1_struct_0(X0) => (v7_struct_0(X0) <=> ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => X1 = X2)))) [input d10_struct_0]
35. ! [X0] : (l1_orders_2(X0) => m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0))) [input dt_k3_yellow_0]
36. ! [X0] : (l1_orders_2(X0) => m1_subset_1(k4_yellow_0(X0),u1_struct_0(X0))) [input dt_k4_yellow_0]
37. ! [X0] : (l1_orders_2(X0) => l1_struct_0(X0)) [input dt_l1_orders_2]
60. ! [X0] : ((l1_orders_2(X0) & v5_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ((r1_orders_2(X0,X2,X1) & r1_orders_2(X0,X1,X2)) => X1 = X2)))) [input t2_orders_2]
62. ! [X0] : ((l1_orders_2(X0) & v1_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => r1_orders_2(X0,k3_yellow_0(X0),X1))) [input t44_yellow_0]
63. ! [X0] : ((l1_orders_2(X0) & v2_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => r1_orders_2(X0,X1,k4_yellow_0(X0)))) [input t45_yellow_0]
77. ? [X0] : ((v7_struct_0(X0) <~> k4_yellow_0(X0) = k3_yellow_0(X0)) & (l1_orders_2(X0) & v3_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
78. ? [X0] : ((v7_struct_0(X0) <~> k4_yellow_0(X0) = k3_yellow_0(X0)) & l1_orders_2(X0) & v3_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) [flattening 77]
82. ! [X0] : (l1_struct_0(X0) | ~l1_orders_2(X0)) [ennf transformation 37]
83. ! [X0] : (m1_subset_1(k4_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0)) [ennf transformation 36]
84. ! [X0] : (m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0)) [ennf transformation 35]
87. ! [X0] : (((v2_yellow_0(X0) & v1_yellow_0(X0)) | ~v3_yellow_0(X0)) | ~l1_orders_2(X0)) [ennf transformation 17]
88. ! [X0] : ((v2_yellow_0(X0) & v1_yellow_0(X0)) | ~v3_yellow_0(X0) | ~l1_orders_2(X0)) [flattening 87]
113. ! [X0] : ((v7_struct_0(X0) <=> ! [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0)))) | ~l1_struct_0(X0)) [ennf transformation 28]
122. ! [X0] : (! [X1] : (r1_orders_2(X0,X1,k4_yellow_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v2_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 63]
123. ! [X0] : (! [X1] : (r1_orders_2(X0,X1,k4_yellow_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v2_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0)) [flattening 122]
124. ! [X0] : (! [X1] : (r1_orders_2(X0,k3_yellow_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 62]
125. ! [X0] : (! [X1] : (r1_orders_2(X0,k3_yellow_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0)) [flattening 124]
128. ! [X0] : (! [X1] : (! [X2] : ((X1 = X2 | (~r1_orders_2(X0,X2,X1) | ~r1_orders_2(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v5_orders_2(X0))) [ennf transformation 60]
129. ! [X0] : (! [X1] : (! [X2] : (X1 = X2 | ~r1_orders_2(X0,X2,X1) | ~r1_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v5_orders_2(X0)) [flattening 128]
150. ? [X0] : (((k4_yellow_0(X0) != k3_yellow_0(X0) | ~v7_struct_0(X0)) & (k4_yellow_0(X0) = k3_yellow_0(X0) | v7_struct_0(X0))) & l1_orders_2(X0) & v3_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) [nnf transformation 78]
151. ? [X0] : ((k4_yellow_0(X0) != k3_yellow_0(X0) | ~v7_struct_0(X0)) & (k4_yellow_0(X0) = k3_yellow_0(X0) | v7_struct_0(X0)) & l1_orders_2(X0) & v3_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) [flattening 150]
152. ? [X0] : ((k4_yellow_0(X0) != k3_yellow_0(X0) | ~v7_struct_0(X0)) & (k4_yellow_0(X0) = k3_yellow_0(X0) | v7_struct_0(X0)) & l1_orders_2(X0) & v3_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) => ((k4_yellow_0(sK0) != k3_yellow_0(sK0) | ~v7_struct_0(sK0)) & (k4_yellow_0(sK0) = k3_yellow_0(sK0) | v7_struct_0(sK0)) & l1_orders_2(sK0) & v3_yellow_0(sK0) & v5_orders_2(sK0) & ~v2_struct_0(sK0)) [choice axiom]
153. (k4_yellow_0(sK0) != k3_yellow_0(sK0) | ~v7_struct_0(sK0)) & (k4_yellow_0(sK0) = k3_yellow_0(sK0) | v7_struct_0(sK0)) & l1_orders_2(sK0) & v3_yellow_0(sK0) & v5_orders_2(sK0) & ~v2_struct_0(sK0) [skolemisation 151,152]
154. ! [X0] : (((v7_struct_0(X0) | ? [X1] : (? [X2] : (X1 != X2 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0)))) & (! [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~v7_struct_0(X0))) | ~l1_struct_0(X0)) [nnf transformation 113]
155. ! [X0] : (((v7_struct_0(X0) | ? [X1] : (? [X2] : (X1 != X2 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0)))) & (! [X3] : (! [X4] : (X3 = X4 | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~v7_struct_0(X0))) | ~l1_struct_0(X0)) [rectify 154]
156. ! [X0] : (? [X1] : (? [X2] : (X1 != X2 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) => (? [X2] : (sK1(X0) != X2 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(sK1(X0),u1_struct_0(X0)))) [choice axiom]
157. ! [X0] : (? [X2] : (sK1(X0) != X2 & m1_subset_1(X2,u1_struct_0(X0))) => (sK1(X0) != sK2(X0) & m1_subset_1(sK2(X0),u1_struct_0(X0)))) [choice axiom]
158. ! [X0] : (((v7_struct_0(X0) | ((sK1(X0) != sK2(X0) & m1_subset_1(sK2(X0),u1_struct_0(X0))) & m1_subset_1(sK1(X0),u1_struct_0(X0)))) & (! [X3] : (! [X4] : (X3 = X4 | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~v7_struct_0(X0))) | ~l1_struct_0(X0)) [skolemisation 155,157,156]
177. ~v2_struct_0(sK0) [cnf transformation 153] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
178. v5_orders_2(sK0) [cnf transformation 153] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
179. v3_yellow_0(sK0) [cnf transformation 153] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
180. l1_orders_2(sK0) [cnf transformation 153] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
181. k4_yellow_0(sK0) = k3_yellow_0(sK0) | v7_struct_0(sK0) [cnf transformation 153] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
182. k4_yellow_0(sK0) != k3_yellow_0(sK0) | ~v7_struct_0(sK0) [cnf transformation 153] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
190. ~l1_orders_2(X0) | l1_struct_0(X0) [cnf transformation 82] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
191. m1_subset_1(k4_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 83] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
192. m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 84] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
195. ~l1_orders_2(X0) | ~v3_yellow_0(X0) | v1_yellow_0(X0) [cnf transformation 88] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
196. ~l1_orders_2(X0) | ~v3_yellow_0(X0) | v2_yellow_0(X0) [cnf transformation 88] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
212. ~m1_subset_1(X4,u1_struct_0(X0)) | X3 = X4 | ~m1_subset_1(X3,u1_struct_0(X0)) | ~v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 158] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
213. m1_subset_1(sK1(X0),u1_struct_0(X0)) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 158] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
214. m1_subset_1(sK2(X0),u1_struct_0(X0)) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 158] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
215. sK1(X0) != sK2(X0) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 158] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
222. r1_orders_2(X0,X1,k4_yellow_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0) [cnf transformation 123] {a:0,w:17,wCS:6800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
223. r1_orders_2(X0,k3_yellow_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0) [cnf transformation 125] {a:0,w:17,wCS:6800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
225. ~r1_orders_2(X0,X2,X1) | X1 = X2 | ~r1_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v5_orders_2(X0) [cnf transformation 129] {a:0,w:23,wCS:9200,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
253. 1 <=> v7_struct_0(sK0) [avatar definition]
255. ~v7_struct_0(sK0) <- (~1) [avatar component clause 253] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
257. 2 <=> k4_yellow_0(sK0) = k3_yellow_0(sK0) [avatar definition]
258. k4_yellow_0(sK0) = k3_yellow_0(sK0) <- (2) [avatar component clause 257] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
259. k4_yellow_0(sK0) != k3_yellow_0(sK0) <- (~2) [avatar component clause 257] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
260. ~1 | ~2 [avatar split clause 182,257,253]
261. 1 | 2 [avatar split clause 181,257,253]
263. 3 <=> l1_orders_2(sK0) [avatar definition]
265. l1_orders_2(sK0) <- (3) [avatar component clause 263] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
266. 3 [avatar split clause 180,263]
268. 4 <=> v3_yellow_0(sK0) [avatar definition]
270. v3_yellow_0(sK0) <- (4) [avatar component clause 268] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
271. 4 [avatar split clause 179,268]
273. 5 <=> v5_orders_2(sK0) [avatar definition]
275. v5_orders_2(sK0) <- (5) [avatar component clause 273] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
276. 5 [avatar split clause 178,273]
278. 6 <=> v2_struct_0(sK0) [avatar definition]
280. ~v2_struct_0(sK0) <- (~6) [avatar component clause 278] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
281. ~6 [avatar split clause 177,278]
363. l1_struct_0(sK0) <- (3) [resolution 190,265] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
366. 23 <=> l1_struct_0(sK0) [avatar definition]
368. l1_struct_0(sK0) <- (23) [avatar component clause 366] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
369. 23 | ~3 [avatar split clause 363,263,366]
414. ~v3_yellow_0(sK0) | v1_yellow_0(sK0) <- (3) [resolution 195,265] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
416. v1_yellow_0(sK0) <- (3, 4) [subsumption resolution 414,270] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
418. 31 <=> v1_yellow_0(sK0) [avatar definition]
420. v1_yellow_0(sK0) <- (31) [avatar component clause 418] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
421. 31 | ~3 | ~4 [avatar split clause 416,268,263,418]
431. ~v3_yellow_0(sK0) | v2_yellow_0(sK0) <- (3) [resolution 196,265] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
433. v2_yellow_0(sK0) <- (3, 4) [subsumption resolution 431,270] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
435. 34 <=> v2_yellow_0(sK0) [avatar definition]
437. v2_yellow_0(sK0) <- (34) [avatar component clause 435] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
438. 34 | ~3 | ~4 [avatar split clause 433,268,263,435]
563. 42 <=> m1_subset_1(k4_yellow_0(sK0),u1_struct_0(sK0)) [avatar definition]
564. m1_subset_1(k4_yellow_0(sK0),u1_struct_0(sK0)) <- (42) [avatar component clause 563] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
565. ~m1_subset_1(k4_yellow_0(sK0),u1_struct_0(sK0)) <- (~42) [avatar component clause 563] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
567. 43 <=> ! [X4] : (k4_yellow_0(sK0) = X4 | ~m1_subset_1(X4,u1_struct_0(sK0)) | ~r1_orders_2(sK0,X4,k4_yellow_0(sK0))) [avatar definition]
568. ~r1_orders_2(sK0,X4,k4_yellow_0(sK0)) | ~m1_subset_1(X4,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X4 <- (43) [avatar component clause 567] {a:2,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
573. ~l1_orders_2(sK0) <- (~42) [resolution 565,191] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
575. $false <- (3, ~42) [subsumption resolution 573,265] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
576. ~3 | 42 [avatar contradiction clause 575]
583. k4_yellow_0(sK0) = X0 | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (42) [resolution 564,212] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
643. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (43) [resolution 568,222] {a:3,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
645. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | ~l1_orders_2(sK0) | ~v2_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (43) [duplicate literal removal 643] {a:3,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
647. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | ~v2_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (3, 43) [subsumption resolution 645,265] {a:3,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
648. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (3, 34, 43) [subsumption resolution 647,437] {a:3,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
649. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | v2_struct_0(sK0) <- (3, 5, 34, 43) [subsumption resolution 648,275] {a:3,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
650. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 <- (3, 5, ~6, 34, 43) [subsumption resolution 649,280] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
658. k4_yellow_0(sK0) = sK1(sK0) | v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (3, 5, ~6, 34, 43) [resolution 650,213] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
659. k4_yellow_0(sK0) = sK2(sK0) | v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (3, 5, ~6, 34, 43) [resolution 650,214] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
663. k4_yellow_0(sK0) = sK1(sK0) | ~l1_struct_0(sK0) <- (~1, 3, 5, ~6, 34, 43) [subsumption resolution 658,255] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
664. k4_yellow_0(sK0) = sK1(sK0) <- (~1, 3, 5, ~6, 23, 34, 43) [subsumption resolution 663,368] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
666. 49 <=> k4_yellow_0(sK0) = sK1(sK0) [avatar definition]
668. k4_yellow_0(sK0) = sK1(sK0) <- (49) [avatar component clause 666] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
669. 49 | 1 | ~3 | ~5 | 6 | ~23 | ~34 | ~43 [avatar split clause 664,567,435,366,278,273,263,253,666]
670. k4_yellow_0(sK0) = sK2(sK0) | ~l1_struct_0(sK0) <- (~1, 3, 5, ~6, 34, 43) [subsumption resolution 659,255] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
671. k4_yellow_0(sK0) = sK2(sK0) <- (~1, 3, 5, ~6, 23, 34, 43) [subsumption resolution 670,368] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
673. 50 <=> k4_yellow_0(sK0) = sK2(sK0) [avatar definition]
675. k4_yellow_0(sK0) = sK2(sK0) <- (50) [avatar component clause 673] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
676. 50 | 1 | ~3 | ~5 | 6 | ~23 | ~34 | ~43 [avatar split clause 671,567,435,366,278,273,263,253,673]
682. k4_yellow_0(sK0) = X0 | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v7_struct_0(sK0) <- (23, 42) [subsumption resolution 583,368] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
684. 52 <=> ! [X0] : (k4_yellow_0(sK0) = X0 | ~m1_subset_1(X0,u1_struct_0(sK0))) [avatar definition]
685. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 <- (52) [avatar component clause 684] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
686. ~1 | 52 | ~23 | ~42 [avatar split clause 682,563,366,684,253]
693. r1_orders_2(sK0,k4_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (2) [superposition 223,258] {a:1,w:17,wCS:1700,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
695. r1_orders_2(sK0,k4_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (2, 3) [subsumption resolution 693,265] {a:1,w:15,wCS:1500,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
696. r1_orders_2(sK0,k4_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (2, 3, 31) [subsumption resolution 695,420] {a:1,w:13,wCS:1300,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
697. r1_orders_2(sK0,k4_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (2, 3, 5, 31) [subsumption resolution 696,275] {a:1,w:11,wCS:1100,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
698. r1_orders_2(sK0,k4_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (2, 3, 5, ~6, 31) [subsumption resolution 697,280] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
700. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | ~r1_orders_2(sK0,X0,k4_yellow_0(sK0)) | ~m1_subset_1(k4_yellow_0(sK0),u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v5_orders_2(sK0) <- (2, 3, 5, ~6, 31) [resolution 698,225] {a:2,w:26,wCS:2600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
701. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | ~r1_orders_2(sK0,X0,k4_yellow_0(sK0)) | ~m1_subset_1(k4_yellow_0(sK0),u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v5_orders_2(sK0) <- (2, 3, 5, ~6, 31) [duplicate literal removal 700] {a:2,w:22,wCS:2200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
702. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | ~r1_orders_2(sK0,X0,k4_yellow_0(sK0)) | ~l1_orders_2(sK0) | ~v5_orders_2(sK0) <- (2, 3, 5, ~6, 31, 42) [subsumption resolution 701,564] {a:2,w:17,wCS:1700,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
703. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | ~r1_orders_2(sK0,X0,k4_yellow_0(sK0)) | ~v5_orders_2(sK0) <- (2, 3, 5, ~6, 31, 42) [subsumption resolution 702,265] {a:2,w:15,wCS:1500,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
704. ~m1_subset_1(X0,u1_struct_0(sK0)) | k4_yellow_0(sK0) = X0 | ~r1_orders_2(sK0,X0,k4_yellow_0(sK0)) <- (2, 3, 5, ~6, 31, 42) [subsumption resolution 703,275] {a:2,w:13,wCS:1300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
705. 43 | ~2 | ~3 | ~5 | 6 | ~31 | ~42 [avatar split clause 704,563,418,278,273,263,257,567]
706. k4_yellow_0(sK0) != sK1(sK0) | v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (50) [superposition 215,675] {a:5,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
708. v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (49, 50) [subsumption resolution 706,668] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
712. v7_struct_0(sK0) <- (23, 49, 50) [subsumption resolution 708,368] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
713. 1 | ~23 | ~49 | ~50 [avatar split clause 712,673,666,366,253]
738. k4_yellow_0(sK0) = k3_yellow_0(sK0) | ~l1_orders_2(sK0) <- (52) [resolution 685,192] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
744. ~l1_orders_2(sK0) <- (~2, 52) [subsumption resolution 738,259] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
745. $false <- (~2, 3, 52) [subsumption resolution 744,265] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
746. 2 | ~3 | ~52 [avatar contradiction clause 745]
749. $false [avatar sat refutation 260,261,266,271,276,281,369,421,438,576,669,676,686,705,713,746] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:327,thDist:-327}
% SZS output end Proof for waybel_7__t3_waybel_7
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 65
% >>> Preprocessing
% Introduced skolems: 12
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 75
% Generated clauses: 365
% Active clauses: 121
% Passive clauses: 147
% Final active clauses: 111
% Final passive clauses: 16
% Main loop iterations started: 127
% 
% >>> Simplifying Inferences
% Duplicate literals: 43
% Fw subsumption resolutions: 43
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 23
% Equational tautologies: 5
% Forward subsumptions: 59
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 115
% Unit resulting resolution: 1
% Forward superposition: 19
% Backward superposition: 15
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 15
% Split components: 30
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 58
% SAT solver unit clauses: 20
% SAT solver binary clauses: 14
% 
% Memory used [KB]: 5245
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption: 0.001 s
% resolution: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

