i: [1,52,0,3,1,1,0,0] t20_zfmisc_1
a: [3,88,0,3,1,52]
s: 88
i: [1,51,0,6,1,1,0,0] t20_zfmisc_1
a: [3,93,0,6,1,51]
s: 93
i: [1,53,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,98,0,2,1,53]
s: 98
i: [1,54,0,2,1,0,0,255] fc1_xboole_0
a: [3,103,0,2,1,54]
s: 103
i: [1,75,0,2,1,0,0,2] rc2_xboole_0
a: [3,113,0,2,1,75]
s: 113
i: [1,76,0,2,1,0,0,2] rc1_xboole_0
a: [3,118,0,2,1,76]
s: 118
i: [1,56,0,3,1,0,0,1] fc2_xboole_0
s: 56
i: [1,71,0,3,1,0,0,1] reflexivity_r1_tarski
s: 71
i: [1,55,0,3,1,0,0,255] d2_xboole_0
a: [3,108,0,3,1,55]
s: 108
r: 98
i: [1,72,0,4,1,0,0,1] fc3_xboole_0
s: 72
i: [1,78,0,4,1,0,0,1] d1_tarski
s: 78
i: [1,57,0,5,2,0,0,255] t6_boole
s: 57
i: [1,64,0,5,2,0,0,2] t7_boole
s: 64
i: [1,81,0,5,1,0,0,1] d2_tarski
s: 81
i: [1,83,0,5,1,0,0,1] d2_tarski
s: 83
i: [1,74,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 74
d: [2,126,1,4,1,0,40,74,78]
s: 126
i: [1,63,0,7,3,0,0,1] t8_boole
s: 63
d: [2,127,1,5,1,0,40,74,81]
s: 127
i: [1,73,0,7,1,0,0,1] commutativity_k2_tarski
s: 73
i: [1,79,0,7,2,0,0,1] d1_tarski
s: 79
d: [2,128,1,5,1,0,40,74,83]
s: 128
i: [1,58,0,9,3,0,0,1] d3_tarski
s: 58
d: [2,129,2,5,1,0,44,127,73]
s: 129
i: [1,84,0,11,3,0,0,1] d2_tarski
s: 84
i: [1,61,0,14,3,0,0,1] d1_tarski
s: 61
d: [2,133,2,5,1,0,44,127,73]
s: 133
i: [1,62,0,14,3,0,0,1] d1_tarski
s: 62
d: [2,121,1,7,3,0,44,57,57]
s: 121
i: [1,69,0,17,3,0,0,1] d2_tarski
s: 69
i: [1,70,0,17,3,0,0,1] d2_tarski
s: 70
d: [2,140,1,7,2,0,40,58,78]
s: 140
i: [1,68,0,23,4,0,0,1] d2_tarski
s: 68
d: [2,141,1,8,2,0,40,58,81]
s: 141
d: [2,142,1,8,2,0,40,58,83]
s: 142
d: [2,149,1,11,3,0,40,61,64]
s: 149
d: [2,158,2,4,1,1,40,141,93]
a: [3,165,2,4,1,158]
s: 165
d: [2,148,1,14,3,0,40,61,74]
s: 148
d: [2,180,3,3,1,1,40,165,79]
a: [3,187,3,3,1,180]
s: 187
r: 165
r: 93
d: [2,196,0,6,1,2,23,93,187]
a: [3,201,0,6,1,196]
s: 201
d: [2,147,1,17,4,0,40,61,58]
s: 147
d: [2,197,4,3,1,2,44,88,187]
a: [3,206,4,3,1,197]
s: 206
d: [2,150,1,17,3,0,40,61,79]
s: 150
d: [2,167,2,4,1,1,40,142,93]
a: [3,174,2,4,1,167]
s: 174
d: [2,155,1,19,4,0,40,68,64]
s: 155
d: [2,154,1,23,4,0,40,68,74]
s: 154
d: [2,208,2,4,1,1,40,201,142]
a: [3,213,2,4,1,208]
s: 213
d: [2,246,5,0,0,2,52,206,213,79]
e: 246
d: [2,247,3,3,1,1,40,213,79]
d: [2,252,3,0,0,2,21,247,206]
e: 252
d: [2,235,2,4,1,2,22,174,187]
f: 252,235,197,180,167,158,51,52
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t20_zfmisc_1
% SZS output start Proof for zfmisc_1__t20_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),k1_tarski(X2)) => X0 = X2) [input t20_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),k1_tarski(X2)) => X0 = X2) [negated conjecture 1]
5. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
6. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
8. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
22. ~! [X0,X1,X2] : (r1_tarski(k2_tarski(X0,X1),k1_tarski(X2)) => X0 = X2) [flattening 2]
23. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 8]
24. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 5]
27. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 6]
29. ! [X0,X1] : (r1_tarski(X0,X1) => ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [unused predicate definition removal 23]
30. ? [X0,X1,X2] : (X0 != X2 & r1_tarski(k2_tarski(X0,X1),k1_tarski(X2))) [ennf transformation 22]
32. ! [X0,X1] : (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1)) [ennf transformation 29]
36. ? [X0,X1,X2] : (X0 != X2 & r1_tarski(k2_tarski(X0,X1),k1_tarski(X2))) => (sK0 != sK2 & r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2))) [choice axiom]
37. sK0 != sK2 & r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2)) [skolemisation 30,36]
38. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 24]
39. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 38]
40. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
41. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 39,40]
42. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 27]
43. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 42]
44. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 43]
45. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
46. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 44,45]
51. r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2)) [cnf transformation 37] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. sK0 != sK2 [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. ~r2_hidden(X2,X0) | r2_hidden(X2,X1) | ~r1_tarski(X0,X1) [cnf transformation 32] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 41] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
66. r2_hidden(X4,X2) | X0 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 46] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
67. r2_hidden(X4,X2) | X1 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 46] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
79. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 59] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
80. r2_hidden(X4,X2) | k2_tarski(X0,X4) != X2 [equality resolution 67] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
81. r2_hidden(X4,k2_tarski(X0,X4)) [equality resolution 80] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
82. r2_hidden(X4,X2) | k2_tarski(X4,X1) != X2 [equality resolution 66] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
83. r2_hidden(X4,k2_tarski(X4,X1)) [equality resolution 82] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
86. 1 <=> sK0 = sK2 [avatar definition]
88. sK0 != sK2 <- (~1) [avatar component clause 86] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
89. ~1 [avatar split clause 52,86]
91. 2 <=> r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2)) [avatar definition]
93. r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2)) <- (2) [avatar component clause 91] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
94. 2 [avatar split clause 51,91]
141. ~r1_tarski(k2_tarski(X4,X2),X3) | r2_hidden(X2,X3) [resolution 58,81] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
142. ~r1_tarski(k2_tarski(X5,X7),X6) | r2_hidden(X5,X6) [resolution 58,83] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
158. r2_hidden(sK1,k1_tarski(sK2)) <- (2) [resolution 141,93] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
163. 8 <=> r2_hidden(sK1,k1_tarski(sK2)) [avatar definition]
165. r2_hidden(sK1,k1_tarski(sK2)) <- (8) [avatar component clause 163] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
166. 8 | ~2 [avatar split clause 158,91,163]
167. r2_hidden(sK0,k1_tarski(sK2)) <- (2) [resolution 142,93] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
172. 9 <=> r2_hidden(sK0,k1_tarski(sK2)) [avatar definition]
174. r2_hidden(sK0,k1_tarski(sK2)) <- (9) [avatar component clause 172] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
175. 9 | ~2 [avatar split clause 167,91,172]
180. sK1 = sK2 <- (8) [resolution 165,79] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
185. 10 <=> sK1 = sK2 [avatar definition]
187. sK1 = sK2 <- (10) [avatar component clause 185] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
188. 10 | ~8 [avatar split clause 180,163,185]
197. sK0 != sK1 <- (~1, 10) [superposition 88,187] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
204. 13 <=> sK0 = sK1 [avatar definition]
206. sK0 != sK1 <- (~13) [avatar component clause 204] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
207. ~13 | 1 | ~10 [avatar split clause 197,185,86,204]
211. 14 <=> r2_hidden(sK0,k1_tarski(sK1)) [avatar definition]
213. r2_hidden(sK0,k1_tarski(sK1)) <- (14) [avatar component clause 211] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
235. r2_hidden(sK0,k1_tarski(sK1)) <- (9, 10) [forward demodulation 174,187] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
236. 14 | ~9 | ~10 [avatar split clause 235,185,172,211]
247. sK0 = sK1 <- (14) [resolution 213,79] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
252. $false <- (~13, 14) [subsumption resolution 247,206] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
253. 13 | ~14 [avatar contradiction clause 252]
254. $false [avatar sat refutation 89,94,166,175,188,207,236,253] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:104,thDist:-104}
% SZS output end Proof for zfmisc_1__t20_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% >>> Preprocessing
% Introduced skolems: 7
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 130
% Active clauses: 43
% Passive clauses: 65
% Final active clauses: 40
% Final passive clauses: 18
% Main loop iterations started: 47
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Trivial inequalities: 2
% Fw subsumption resolutions: 2
% Fw demodulations: 1
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Equational tautologies: 3
% Forward subsumptions: 33
% 
% >>> Generating Inferences
% Binary resolution: 44
% Unit resulting resolution: 1
% Forward superposition: 8
% Backward superposition: 18
% Self superposition: 1
% Equality factoring: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 19
% SAT solver unit clauses: 7
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

