i: [1,84,0,2,1,1,0,0] t33_mssubfam
a: [3,171,0,2,1,84]
s: 171
i: [1,82,0,2,1,1,0,0] t33_mssubfam
a: [3,181,0,2,1,82]
s: 181
i: [1,80,0,2,1,1,0,0] t33_mssubfam
a: [3,191,0,2,1,80]
s: 191
i: [1,78,0,2,1,1,0,0] t33_mssubfam
a: [3,201,0,2,1,78]
s: 201
i: [1,85,0,3,1,1,0,0] t33_mssubfam
a: [3,166,0,3,1,85]
s: 166
i: [1,83,0,3,1,1,0,0] t33_mssubfam
a: [3,176,0,3,1,83]
s: 176
i: [1,81,0,3,1,1,0,0] t33_mssubfam
a: [3,186,0,3,1,81]
s: 186
i: [1,79,0,3,1,1,0,0] t33_mssubfam
a: [3,196,0,3,1,79]
s: 196
i: [1,88,0,4,1,1,0,0] t33_mssubfam
a: [3,151,0,4,1,88]
s: 151
i: [1,87,0,4,1,1,0,0] t33_mssubfam
a: [3,156,0,4,1,87]
s: 156
i: [1,86,0,6,1,1,0,0] t33_mssubfam
a: [3,161,0,6,1,86]
s: 161
i: [1,142,0,2,1,0,0,2] rc2_relat_1
a: [3,206,0,2,1,142]
s: 206
i: [1,141,0,2,1,0,0,2] rc2_relat_1
a: [3,211,0,2,1,141]
s: 211
i: [1,145,0,2,1,0,0,2] rc3_funct_1
a: [3,216,0,2,1,145]
s: 216
i: [1,144,0,2,1,0,0,2] rc3_funct_1
a: [3,221,0,2,1,144]
s: 221
i: [1,143,0,2,1,0,0,2] rc3_funct_1
a: [3,226,0,2,1,143]
s: 226
i: [1,147,0,2,1,0,0,1] rc1_funct_1
a: [3,231,0,2,1,147]
s: 231
i: [1,146,0,2,1,0,0,1] rc1_funct_1
a: [3,236,0,2,1,146]
s: 236
i: [1,94,0,3,1,0,0,2] rc9_funct_1
s: 94
i: [1,95,0,3,1,0,0,2] rc9_funct_1
s: 95
i: [1,97,0,3,1,0,0,2] rc9_funct_1
s: 97
i: [1,98,0,3,1,0,0,2] rc4_funcop_1
s: 98
i: [1,99,0,3,1,0,0,2] rc4_funcop_1
s: 99
i: [1,101,0,3,1,0,0,2] rc4_funcop_1
s: 101
i: [1,103,0,3,1,0,0,2] rc3_pboole
s: 103
i: [1,104,0,3,1,0,0,2] rc3_pboole
s: 104
i: [1,106,0,3,1,0,0,2] rc3_pboole
s: 106
i: [1,108,0,3,1,0,0,1] rc1_pboole
s: 108
i: [1,110,0,3,1,0,0,1] rc1_pboole
s: 110
i: [1,96,0,4,1,0,0,2] rc9_funct_1
s: 96
i: [1,100,0,4,1,0,0,2] rc4_funcop_1
s: 100
i: [1,102,0,4,1,0,0,2] rc4_funcop_1
s: 102
i: [1,105,0,4,1,0,0,2] rc3_pboole
s: 105
i: [1,107,0,4,1,0,0,2] rc3_pboole
s: 107
i: [1,109,0,4,1,0,0,1] rc1_pboole
s: 109
i: [1,111,0,4,1,0,0,1] rc1_pboole
s: 111
i: [1,120,0,14,5,0,0,1] fc1_mboolean
s: 120
i: [1,121,0,14,5,0,0,1] fc1_mboolean
s: 121
i: [1,123,0,14,5,0,0,1] fc1_mboolean
s: 123
i: [1,125,0,14,5,0,0,1] dt_k1_mboolean
s: 125
i: [1,127,0,14,5,0,0,1] dt_k1_mboolean
s: 127
i: [1,135,0,14,5,0,0,1] rc2_mssubfam
s: 135
i: [1,136,0,14,5,0,0,1] rc2_mssubfam
s: 136
i: [1,138,0,14,5,0,0,1] rc2_mssubfam
s: 138
i: [1,122,0,15,5,0,0,1] fc1_mboolean
s: 122
i: [1,124,0,15,5,0,0,1] fc1_mboolean
s: 124
i: [1,126,0,15,5,0,0,1] dt_k1_mboolean
s: 126
i: [1,128,0,15,5,0,0,1] dt_k1_mboolean
s: 128
i: [1,137,0,15,5,0,0,1] rc2_mssubfam
s: 137
i: [1,139,0,15,5,0,0,1] rc2_mssubfam
s: 139
i: [1,113,0,16,6,0,0,1] rc5_pboole
s: 113
i: [1,114,0,16,6,0,0,1] rc5_pboole
s: 114
i: [1,116,0,16,6,0,0,1] rc5_pboole
s: 116
i: [1,129,0,16,6,0,0,1] dt_m3_pboole
s: 129
i: [1,131,0,16,6,0,0,1] dt_m3_pboole
s: 131
i: [1,133,0,16,5,0,0,1] existence_m3_pboole
s: 133
d: [2,238,1,20,5,1,40,129,161]
a: [3,242,1,4,1,238]
s: 242
d: [2,270,2,0,0,5,52,181,171,176,166,242,120]
e: 270
d: [2,271,2,10,4,1,40,242,120]
d: [2,273,2,7,3,2,21,271,166]
d: [2,274,2,5,2,3,21,273,171]
d: [2,275,2,2,1,4,21,274,176]
d: [2,276,2,0,0,5,21,275,181]
e: 276
r: 242
i: [1,115,0,17,6,0,0,1] rc5_pboole
s: 115
a: [3,241,1,4,1,238]
s: 241
i: [1,117,0,17,6,0,0,1] rc5_pboole
s: 117
i: [1,118,0,17,5,0,0,255] redefinition_k1_mssubfam
s: 118
a: [3,246,1,5,1,238]
s: 246
d: [2,278,2,0,0,5,52,181,171,176,166,246,122]
e: 278
d: [2,279,2,10,4,1,40,246,122]
d: [2,281,2,7,3,2,21,279,166]
d: [2,282,2,5,2,3,21,281,171]
d: [2,283,2,2,1,4,21,282,176]
d: [2,284,2,0,0,5,21,283,181]
e: 284
r: 246
i: [1,130,0,17,6,0,0,1] dt_m3_pboole
s: 130
a: [3,250,1,4,1,238]
s: 250
d: [2,295,2,0,0,5,52,181,171,176,166,250,123]
e: 295
d: [2,296,2,10,4,1,40,250,123]
d: [2,298,2,7,3,2,21,296,166]
d: [2,299,2,5,2,3,21,298,171]
d: [2,300,2,2,1,4,21,299,176]
d: [2,301,2,0,0,5,21,300,181]
e: 301
r: 250
i: [1,132,0,17,6,0,0,1] dt_m3_pboole
s: 132
i: [1,112,0,18,6,0,0,1] rc5_pboole
s: 112
a: [3,249,1,4,1,238]
s: 249
i: [1,119,0,18,5,0,0,255] dt_k1_mssubfam
s: 119
a: [3,245,1,5,1,238]
s: 245
i: [1,134,0,18,5,0,0,1] rc2_mssubfam
s: 134
i: [1,140,0,24,9,0,0,2] reflexivity_r2_pboole
s: 140
a: [3,254,1,5,1,238]
s: 254
d: [2,346,2,0,0,5,52,181,171,176,166,254,124]
e: 346
d: [2,347,2,10,4,1,40,254,124]
d: [2,349,2,7,3,2,21,347,166]
d: [2,350,2,5,2,3,21,349,171]
d: [2,351,2,2,1,4,21,350,176]
d: [2,352,2,0,0,5,21,351,181]
e: 352
r: 254
i: [1,91,0,28,10,0,0,2] d18_pboole
s: 91
a: [3,258,1,2,1,238]
s: 258
d: [2,354,0,26,9,0,21,91,129]
d: [2,355,0,24,8,0,21,354,131]
d: [2,356,0,21,7,0,21,355,130]
d: [2,357,0,18,6,0,21,356,132]
s: 357
i: [1,92,0,28,10,0,0,2] d18_pboole
s: 92
d: [2,260,1,20,5,1,40,131,161]
a: [3,264,1,2,1,260]
s: 264
i: [1,89,0,30,10,0,0,1] t18_mboolean
s: 89
d: [2,286,1,21,5,1,40,130,161]
d: [2,289,1,17,4,2,21,286,241]
a: [3,293,1,3,1,289]
s: 293
i: [1,90,0,30,10,0,0,1] t18_mboolean
s: 90
i: [1,93,0,42,15,0,0,1] t9_pboole
s: 93
d: [2,303,1,21,5,1,40,132,161]
d: [2,306,1,17,4,2,21,303,241]
a: [3,310,1,3,1,306]
s: 310
a: [3,253,1,5,1,238]
s: 253
d: [2,359,1,24,5,1,40,357,161]
d: [2,369,1,20,4,2,21,359,249]
d: [2,370,1,15,3,3,21,369,245]
d: [2,371,1,11,2,4,21,370,241]
a: [3,375,1,6,1,371]
s: 375
d: [2,267,1,24,9,0,40,133,129]
d: [2,268,1,14,5,0,18,267]
s: 268
d: [2,266,1,24,9,0,40,133,131]
d: [2,269,1,14,5,0,18,266]
s: 269
d: [2,387,2,48,14,1,40,375,93]
d: [2,389,2,43,13,2,21,387,253]
d: [2,390,2,39,12,3,21,389,249]
d: [2,391,2,34,11,4,21,390,245]
d: [2,392,2,30,10,5,21,391,241]
d: [2,393,2,27,9,6,21,392,310]
d: [2,394,2,25,8,7,21,393,264]
d: [2,395,2,22,7,8,21,394,293]
d: [2,396,2,20,6,9,21,395,258]
s: 396
d: [2,322,1,32,9,0,40,119,131]
d: [2,334,1,27,8,0,21,322,124]
d: [2,335,1,23,7,0,21,334,123]
d: [2,336,1,18,6,0,21,335,122]
d: [2,337,1,14,5,0,21,336,120]
s: 337
d: [2,397,3,38,14,9,40,396,90]
d: [2,398,3,28,10,9,18,397]
d: [2,399,3,25,9,10,21,398,166]
d: [2,400,3,23,8,11,21,399,171]
d: [2,401,3,20,7,12,21,400,176]
d: [2,402,3,18,6,13,21,401,181]
s: 402
d: [2,323,1,32,9,0,40,119,129]
d: [2,338,1,27,8,0,21,323,124]
d: [2,339,1,23,7,0,21,338,123]
d: [2,340,1,18,6,0,21,339,122]
d: [2,341,1,14,5,0,21,340,120]
s: 341
d: [2,287,1,25,9,0,40,130,133]
d: [2,288,1,15,5,0,18,287]
s: 288
d: [2,405,4,14,5,14,40,402,156]
d: [2,406,4,11,4,15,21,405,186]
d: [2,407,4,9,3,16,21,406,191]
d: [2,408,4,6,2,17,21,407,196]
d: [2,409,4,4,1,18,21,408,201]
a: [3,413,4,4,1,409]
s: 413
d: [2,417,5,0,0,10,52,181,201,171,191,176,166,196,186,151,413,92]
e: 417
d: [2,419,5,24,9,1,40,413,92]
d: [2,429,5,20,8,2,21,419,151]
d: [2,430,5,17,7,3,21,429,186]
d: [2,431,5,15,6,4,21,430,191]
d: [2,432,5,12,5,5,21,431,196]
d: [2,433,5,10,4,6,21,432,201]
d: [2,434,5,7,3,7,21,433,166]
d: [2,435,5,5,2,8,21,434,171]
d: [2,436,5,2,1,9,21,435,176]
d: [2,437,5,0,0,10,21,436,181]
e: 437
f: 437,409,371,352,306,301,289,284,276,260,238,78,79,80,81,82,83,84,85,86,87,88
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mssubfam__t33_mssubfam
% SZS output start Proof for mssubfam__t33_mssubfam
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : (m3_pboole(X3,X0,k1_mboolean(X0,X2)) => (r1_pboole(X0,X1,X3) => m3_pboole(X1,X0,X2))))) [input t33_mssubfam]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : (m3_pboole(X3,X0,k1_mboolean(X0,X2)) => (r1_pboole(X0,X1,X3) => m3_pboole(X1,X0,X2))))) [negated conjecture 1]
3. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (m3_pboole(X2,X0,X1) <=> r2_pboole(X0,X2,X1)))) [input d18_pboole]
6. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (m3_pboole(X2,X0,X1) => (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)))) [input dt_m3_pboole]
8. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k1_mboolean(X0,X1),X0) & v1_funct_1(k1_mboolean(X0,X1)) & v4_relat_1(k1_mboolean(X0,X1),X0) & v2_relat_1(k1_mboolean(X0,X1)) & v1_relat_1(k1_mboolean(X0,X1)))) [input fc1_mboolean]
20. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r2_pboole(X0,X1,X2) <=> r1_pboole(X0,X1,k1_mboolean(X0,X2))))) [input t18_mboolean]
21. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r1_pboole(X0,X1,X2)) => r1_pboole(X0,X1,X3))))) [input t9_pboole]
22. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : (m3_pboole(X3,X0,k1_mboolean(X0,X2)) => (r1_pboole(X0,X1,X3) => m3_pboole(X1,X0,X2))))) [flattening 2]
23. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r2_pboole(X0,X1,X2) <=> r1_pboole(X0,X1,k1_mboolean(X0,X2))))) [flattening 20]
24. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (m3_pboole(X2,X0,X1) <=> r2_pboole(X0,X2,X1)))) [flattening 3]
25. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r1_pboole(X0,X1,X2)) => r1_pboole(X0,X1,X3))))) [flattening 21]
26. ? [X0,X1] : (? [X2] : (? [X3] : ((~m3_pboole(X1,X0,X2) & r1_pboole(X0,X1,X3)) & m3_pboole(X3,X0,k1_mboolean(X0,X2))) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 22]
27. ? [X0,X1] : (? [X2] : (? [X3] : (~m3_pboole(X1,X0,X2) & r1_pboole(X0,X1,X3) & m3_pboole(X3,X0,k1_mboolean(X0,X2))) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 26]
28. ! [X0,X1] : (! [X2] : ((r2_pboole(X0,X1,X2) <=> r1_pboole(X0,X1,k1_mboolean(X0,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 23]
29. ! [X0,X1] : (! [X2] : ((r2_pboole(X0,X1,X2) <=> r1_pboole(X0,X1,k1_mboolean(X0,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 28]
30. ! [X0,X1] : (! [X2] : ((m3_pboole(X2,X0,X1) <=> r2_pboole(X0,X2,X1)) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 24]
31. ! [X0,X1] : (! [X2] : ((m3_pboole(X2,X0,X1) <=> r2_pboole(X0,X2,X1)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 30]
32. ! [X0,X1] : (! [X2] : (! [X3] : ((r1_pboole(X0,X1,X3) | (~r2_pboole(X0,X2,X3) | ~r1_pboole(X0,X1,X2))) | (~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 25]
33. ! [X0,X1] : (! [X2] : (! [X3] : (r1_pboole(X0,X1,X3) | ~r2_pboole(X0,X2,X3) | ~r1_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 32]
40. ! [X0,X1] : ((v1_partfun1(k1_mboolean(X0,X1),X0) & v1_funct_1(k1_mboolean(X0,X1)) & v4_relat_1(k1_mboolean(X0,X1),X0) & v2_relat_1(k1_mboolean(X0,X1)) & v1_relat_1(k1_mboolean(X0,X1))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 8]
41. ! [X0,X1] : ((v1_partfun1(k1_mboolean(X0,X1),X0) & v1_funct_1(k1_mboolean(X0,X1)) & v4_relat_1(k1_mboolean(X0,X1),X0) & v2_relat_1(k1_mboolean(X0,X1)) & v1_relat_1(k1_mboolean(X0,X1))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 40]
44. ! [X0,X1] : (! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) | ~m3_pboole(X2,X0,X1)) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 6]
45. ! [X0,X1] : (! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) | ~m3_pboole(X2,X0,X1)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 44]
52. ? [X0,X1] : (? [X2] : (? [X3] : (~m3_pboole(X1,X0,X2) & r1_pboole(X0,X1,X3) & m3_pboole(X3,X0,k1_mboolean(X0,X2))) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (~m3_pboole(sK1,sK0,X2) & r1_pboole(sK0,sK1,X3) & m3_pboole(X3,sK0,k1_mboolean(sK0,X2))) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
53. ? [X2] : (? [X3] : (~m3_pboole(sK1,sK0,X2) & r1_pboole(sK0,sK1,X3) & m3_pboole(X3,sK0,k1_mboolean(sK0,X2))) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (? [X3] : (~m3_pboole(sK1,sK0,sK2) & r1_pboole(sK0,sK1,X3) & m3_pboole(X3,sK0,k1_mboolean(sK0,sK2))) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
54. ? [X3] : (~m3_pboole(sK1,sK0,sK2) & r1_pboole(sK0,sK1,X3) & m3_pboole(X3,sK0,k1_mboolean(sK0,sK2))) => (~m3_pboole(sK1,sK0,sK2) & r1_pboole(sK0,sK1,sK3) & m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2))) [choice axiom]
55. ((~m3_pboole(sK1,sK0,sK2) & r1_pboole(sK0,sK1,sK3) & m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2))) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 27,54,53,52]
56. ! [X0,X1] : (! [X2] : (((r2_pboole(X0,X1,X2) | ~r1_pboole(X0,X1,k1_mboolean(X0,X2))) & (r1_pboole(X0,X1,k1_mboolean(X0,X2)) | ~r2_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 29]
57. ! [X0,X1] : (! [X2] : (((m3_pboole(X2,X0,X1) | ~r2_pboole(X0,X2,X1)) & (r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 31]
78. v1_relat_1(sK1) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
79. v4_relat_1(sK1,sK0) [cnf transformation 55] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
80. v1_funct_1(sK1) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
81. v1_partfun1(sK1,sK0) [cnf transformation 55] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
82. v1_relat_1(sK2) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
83. v4_relat_1(sK2,sK0) [cnf transformation 55] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
84. v1_funct_1(sK2) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
85. v1_partfun1(sK2,sK0) [cnf transformation 55] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
86. m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2)) [cnf transformation 55] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
87. r1_pboole(sK0,sK1,sK3) [cnf transformation 55] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
88. ~m3_pboole(sK1,sK0,sK2) [cnf transformation 55] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
90. ~r1_pboole(X0,X1,k1_mboolean(X0,X2)) | r2_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 56] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
91. r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 57] {a:0,w:28,wCS:11200,sine:2,thAx:0,allAx:1,thDist:-1}
92. ~r2_pboole(X0,X2,X1) | m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 57] {a:0,w:28,wCS:11200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
93. ~r2_pboole(X0,X2,X3) | r1_pboole(X0,X1,X3) | ~r1_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 33] {a:0,w:42,wCS:16800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. v1_relat_1(k1_mboolean(X0,X1)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 41] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
122. v4_relat_1(k1_mboolean(X0,X1),X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 41] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
123. v1_funct_1(k1_mboolean(X0,X1)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 41] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
124. v1_partfun1(k1_mboolean(X0,X1),X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 41] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
129. ~m3_pboole(X2,X0,X1) | v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 45] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
130. ~m3_pboole(X2,X0,X1) | v4_relat_1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 45] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
131. ~m3_pboole(X2,X0,X1) | v1_funct_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 45] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
132. ~m3_pboole(X2,X0,X1) | v1_partfun1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 45] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
149. 1 <=> m3_pboole(sK1,sK0,sK2) [avatar definition]
151. ~m3_pboole(sK1,sK0,sK2) <- (~1) [avatar component clause 149] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
152. ~1 [avatar split clause 88,149]
154. 2 <=> r1_pboole(sK0,sK1,sK3) [avatar definition]
156. r1_pboole(sK0,sK1,sK3) <- (2) [avatar component clause 154] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
157. 2 [avatar split clause 87,154]
159. 3 <=> m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2)) [avatar definition]
161. m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2)) <- (3) [avatar component clause 159] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
162. 3 [avatar split clause 86,159]
164. 4 <=> v1_partfun1(sK2,sK0) [avatar definition]
166. v1_partfun1(sK2,sK0) <- (4) [avatar component clause 164] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
167. 4 [avatar split clause 85,164]
169. 5 <=> v1_funct_1(sK2) [avatar definition]
171. v1_funct_1(sK2) <- (5) [avatar component clause 169] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
172. 5 [avatar split clause 84,169]
174. 6 <=> v4_relat_1(sK2,sK0) [avatar definition]
176. v4_relat_1(sK2,sK0) <- (6) [avatar component clause 174] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
177. 6 [avatar split clause 83,174]
179. 7 <=> v1_relat_1(sK2) [avatar definition]
181. v1_relat_1(sK2) <- (7) [avatar component clause 179] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
182. 7 [avatar split clause 82,179]
184. 8 <=> v1_partfun1(sK1,sK0) [avatar definition]
186. v1_partfun1(sK1,sK0) <- (8) [avatar component clause 184] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
187. 8 [avatar split clause 81,184]
189. 9 <=> v1_funct_1(sK1) [avatar definition]
191. v1_funct_1(sK1) <- (9) [avatar component clause 189] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
192. 9 [avatar split clause 80,189]
194. 10 <=> v4_relat_1(sK1,sK0) [avatar definition]
196. v4_relat_1(sK1,sK0) <- (10) [avatar component clause 194] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
197. 10 [avatar split clause 79,194]
199. 11 <=> v1_relat_1(sK1) [avatar definition]
201. v1_relat_1(sK1) <- (11) [avatar component clause 199] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
202. 11 [avatar split clause 78,199]
238. v1_relat_1(sK3) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (3) [resolution 129,161] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
240. 19 <=> v1_relat_1(k1_mboolean(sK0,sK2)) [avatar definition]
241. v1_relat_1(k1_mboolean(sK0,sK2)) <- (19) [avatar component clause 240] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
242. ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (~19) [avatar component clause 240] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
244. 20 <=> v4_relat_1(k1_mboolean(sK0,sK2),sK0) [avatar definition]
245. v4_relat_1(k1_mboolean(sK0,sK2),sK0) <- (20) [avatar component clause 244] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
246. ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) <- (~20) [avatar component clause 244] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
248. 21 <=> v1_funct_1(k1_mboolean(sK0,sK2)) [avatar definition]
249. v1_funct_1(k1_mboolean(sK0,sK2)) <- (21) [avatar component clause 248] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
250. ~v1_funct_1(k1_mboolean(sK0,sK2)) <- (~21) [avatar component clause 248] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
252. 22 <=> v1_partfun1(k1_mboolean(sK0,sK2),sK0) [avatar definition]
253. v1_partfun1(k1_mboolean(sK0,sK2),sK0) <- (22) [avatar component clause 252] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
254. ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) <- (~22) [avatar component clause 252] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
256. 23 <=> v1_relat_1(sK3) [avatar definition]
258. v1_relat_1(sK3) <- (23) [avatar component clause 256] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
259. ~19 | ~20 | ~21 | ~22 | 23 | ~3 [avatar split clause 238,159,256,252,248,244,240]
260. v1_funct_1(sK3) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (3) [resolution 131,161] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
262. 24 <=> v1_funct_1(sK3) [avatar definition]
264. v1_funct_1(sK3) <- (24) [avatar component clause 262] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
265. ~19 | ~20 | ~21 | ~22 | 24 | ~3 [avatar split clause 260,159,262,252,248,244,240]
271. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~19) [resolution 242,120] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
273. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, ~19) [subsumption resolution 271,166] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
274. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, 5, ~19) [subsumption resolution 273,171] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
275. ~v1_relat_1(sK2) <- (4, 5, 6, ~19) [subsumption resolution 274,176] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
276. $false <- (4, 5, 6, 7, ~19) [subsumption resolution 275,181] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
277. ~4 | ~5 | ~6 | ~7 | 19 [avatar contradiction clause 276]
279. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~20) [resolution 246,122] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
281. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, ~20) [subsumption resolution 279,166] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
282. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, 5, ~20) [subsumption resolution 281,171] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
283. ~v1_relat_1(sK2) <- (4, 5, 6, ~20) [subsumption resolution 282,176] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
284. $false <- (4, 5, 6, 7, ~20) [subsumption resolution 283,181] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
285. ~4 | ~5 | ~6 | ~7 | 20 [avatar contradiction clause 284]
286. v4_relat_1(sK3,sK0) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (3) [resolution 130,161] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
289. v4_relat_1(sK3,sK0) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) <- (3, 19) [subsumption resolution 286,241] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
291. 25 <=> v4_relat_1(sK3,sK0) [avatar definition]
293. v4_relat_1(sK3,sK0) <- (25) [avatar component clause 291] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
294. ~20 | ~21 | ~22 | 25 | ~3 | ~19 [avatar split clause 289,240,159,291,252,248,244]
296. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~21) [resolution 250,123] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
298. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, ~21) [subsumption resolution 296,166] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
299. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, 5, ~21) [subsumption resolution 298,171] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
300. ~v1_relat_1(sK2) <- (4, 5, 6, ~21) [subsumption resolution 299,176] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
301. $false <- (4, 5, 6, 7, ~21) [subsumption resolution 300,181] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
302. ~4 | ~5 | ~6 | ~7 | 21 [avatar contradiction clause 301]
303. v1_partfun1(sK3,sK0) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (3) [resolution 132,161] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
306. v1_partfun1(sK3,sK0) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) <- (3, 19) [subsumption resolution 303,241] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
308. 26 <=> v1_partfun1(sK3,sK0) [avatar definition]
310. v1_partfun1(sK3,sK0) <- (26) [avatar component clause 308] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
311. ~20 | ~21 | ~22 | 26 | ~3 | ~19 [avatar split clause 306,240,159,308,252,248,244]
347. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~22) [resolution 254,124] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
349. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, ~22) [subsumption resolution 347,166] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
350. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, 5, ~22) [subsumption resolution 349,171] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
351. ~v1_relat_1(sK2) <- (4, 5, 6, ~22) [subsumption resolution 350,176] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
352. $false <- (4, 5, 6, 7, ~22) [subsumption resolution 351,181] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
353. ~4 | ~5 | ~6 | ~7 | 22 [avatar contradiction clause 352]
354. r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [subsumption resolution 91,129] {a:0,w:26,wCS:10400,sine:1,thAx:0,allAx:1,thDist:-1}
355. r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v4_relat_1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [subsumption resolution 354,131] {a:0,w:24,wCS:9600,sine:1,thAx:0,allAx:1,thDist:-1}
356. r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [subsumption resolution 355,130] {a:0,w:21,wCS:8400,sine:1,thAx:0,allAx:1,thDist:-1}
357. ~m3_pboole(X2,X0,X1) | r2_pboole(X0,X2,X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [subsumption resolution 356,132] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
359. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (3) [resolution 357,161] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
369. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (3, 21) [subsumption resolution 359,249] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
370. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (3, 20, 21) [subsumption resolution 369,245] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
371. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) <- (3, 19, 20, 21) [subsumption resolution 370,241] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
373. 27 <=> r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) [avatar definition]
375. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) <- (27) [avatar component clause 373] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
376. ~22 | 27 | ~3 | ~19 | ~20 | ~21 [avatar split clause 371,248,244,240,159,373,252]
387. r1_pboole(sK0,X0,k1_mboolean(sK0,sK2)) | ~r1_pboole(sK0,X0,sK3) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (27) [resolution 375,93] {a:2,w:48,wCS:4800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
389. r1_pboole(sK0,X0,k1_mboolean(sK0,sK2)) | ~r1_pboole(sK0,X0,sK3) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (22, 27) [subsumption resolution 387,253] {a:2,w:43,wCS:4300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
390. r1_pboole(sK0,X0,k1_mboolean(sK0,sK2)) | ~r1_pboole(sK0,X0,sK3) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (21, 22, 27) [subsumption resolution 389,249] {a:2,w:39,wCS:3900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
391. r1_pboole(sK0,X0,k1_mboolean(sK0,sK2)) | ~r1_pboole(sK0,X0,sK3) | ~v1_relat_1(k1_mboolean(sK0,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (20, 21, 22, 27) [subsumption resolution 390,245] {a:2,w:34,wCS:3400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
392. r1_pboole(sK0,X0,k1_mboolean(sK0,sK2)) | ~r1_pboole(sK0,X0,sK3) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (19, 20, 21, 22, 27) [subsumption resolution 391,241] {a:2,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
393. r1_pboole(sK0,X0,k1_mboolean(sK0,sK2)) | ~r1_pboole(sK0,X0,sK3) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (19, 20, 21, 22, 26, 27) [subsumption resolution 392,310] {a:2,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
394. r1_pboole(sK0,X0,k1_mboolean(sK0,sK2)) | ~r1_pboole(sK0,X0,sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (19, 20, 21, 22, 24, 26, 27) [subsumption resolution 393,264] {a:2,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
395. r1_pboole(sK0,X0,k1_mboolean(sK0,sK2)) | ~r1_pboole(sK0,X0,sK3) | ~v1_relat_1(sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (19, 20, 21, 22, 24, 25, 26, 27) [subsumption resolution 394,293] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
396. r1_pboole(sK0,X0,k1_mboolean(sK0,sK2)) | ~r1_pboole(sK0,X0,sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (19, 20, 21, 22, 23, 24, 25, 26, 27) [subsumption resolution 395,258] {a:2,w:20,wCS:2000,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
397. ~r1_pboole(sK0,X0,sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | r2_pboole(sK0,X0,sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (19, 20, 21, 22, 23, 24, 25, 26, 27) [resolution 396,90] {a:3,w:38,wCS:3800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
398. ~r1_pboole(sK0,X0,sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | r2_pboole(sK0,X0,sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (19, 20, 21, 22, 23, 24, 25, 26, 27) [duplicate literal removal 397] {a:3,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
399. ~r1_pboole(sK0,X0,sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | r2_pboole(sK0,X0,sK2) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, 19, 20, 21, 22, 23, 24, 25, 26, 27) [subsumption resolution 398,166] {a:3,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
400. ~r1_pboole(sK0,X0,sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | r2_pboole(sK0,X0,sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (4, 5, 19, 20, 21, 22, 23, 24, 25, 26, 27) [subsumption resolution 399,171] {a:3,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
401. ~r1_pboole(sK0,X0,sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | r2_pboole(sK0,X0,sK2) | ~v1_relat_1(sK2) <- (4, 5, 6, 19, 20, 21, 22, 23, 24, 25, 26, 27) [subsumption resolution 400,176] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
402. ~r1_pboole(sK0,X0,sK3) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | r2_pboole(sK0,X0,sK2) <- (4, 5, 6, 7, 19, 20, 21, 22, 23, 24, 25, 26, 27) [subsumption resolution 401,181] {a:3,w:18,wCS:1800,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
405. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | r2_pboole(sK0,sK1,sK2) <- (2, 4, 5, 6, 7, 19, 20, 21, 22, 23, 24, 25, 26, 27) [resolution 402,156] {a:4,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
406. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | r2_pboole(sK0,sK1,sK2) <- (2, 4, 5, 6, 7, 8, 19, 20, 21, 22, 23, 24, 25, 26, 27) [subsumption resolution 405,186] {a:4,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
407. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | r2_pboole(sK0,sK1,sK2) <- (2, 4, 5, 6, 7, 8, 9, 19, 20, 21, 22, 23, 24, 25, 26, 27) [subsumption resolution 406,191] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
408. ~v1_relat_1(sK1) | r2_pboole(sK0,sK1,sK2) <- (2, 4, 5, 6, 7, 8, 9, 10, 19, 20, 21, 22, 23, 24, 25, 26, 27) [subsumption resolution 407,196] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
409. r2_pboole(sK0,sK1,sK2) <- (2, 4, 5, 6, 7, 8, 9, 10, 11, 19, 20, 21, 22, 23, 24, 25, 26, 27) [subsumption resolution 408,201] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
411. 28 <=> r2_pboole(sK0,sK1,sK2) [avatar definition]
413. r2_pboole(sK0,sK1,sK2) <- (28) [avatar component clause 411] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
414. 28 | ~2 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~19 | ~20 | ~21 | ~22 | ~23 | ~24 | ~25 | ~26 | ~27 [avatar split clause 409,373,308,291,262,256,252,248,244,240,199,194,189,184,179,174,169,164,154,411]
419. m3_pboole(sK1,sK0,sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (28) [resolution 413,92] {a:5,w:24,wCS:2400,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
429. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 28) [subsumption resolution 419,151] {a:5,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
430. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 8, 28) [subsumption resolution 429,186] {a:5,w:17,wCS:1700,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
431. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 8, 9, 28) [subsumption resolution 430,191] {a:5,w:15,wCS:1500,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
432. ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 8, 9, 10, 28) [subsumption resolution 431,196] {a:5,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
433. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 8, 9, 10, 11, 28) [subsumption resolution 432,201] {a:5,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
434. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 4, 8, 9, 10, 11, 28) [subsumption resolution 433,166] {a:5,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
435. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 4, 5, 8, 9, 10, 11, 28) [subsumption resolution 434,171] {a:5,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
436. ~v1_relat_1(sK2) <- (~1, 4, 5, 6, 8, 9, 10, 11, 28) [subsumption resolution 435,176] {a:5,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
437. $false <- (~1, 4, 5, 6, 7, 8, 9, 10, 11, 28) [subsumption resolution 436,181] {a:5,w:0,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
438. 1 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~28 [avatar contradiction clause 437]
439. $false [avatar sat refutation 152,157,162,167,172,177,182,187,192,197,202,259,265,277,285,294,302,311,353,376,414,438] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:333,thDist:-333}
% SZS output end Proof for mssubfam__t33_mssubfam
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% Introduced skolems: 14
% 
% >>> Saturation
% Initial clauses: 70
% Generated clauses: 247
% Active clauses: 87
% Passive clauses: 101
% Final active clauses: 83
% Final passive clauses: 9
% Main loop iterations started: 92
% 
% >>> Simplifying Inferences
% Duplicate literals: 80
% Fw subsumption resolutions: 82
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 15
% 
% >>> Generating Inferences
% Binary resolution: 38
% Unit resulting resolution: 5
% Forward superposition: 3
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 20
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 36
% SAT solver unit clauses: 18
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% splitting model update: 0.001 s ( own 0.0000 s ) 
% backward demodulation index maintenance: 0.001 s
% forward subsumption: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

