i: [1,13,0,15,1,1,0,0] t113_xboole_1
a: [3,20,0,15,1,13]
s: 20
i: [1,15,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 15
i: [1,16,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 16
r: 20
d: [2,22,0,15,1,1,23,20,16]
d: [2,23,0,15,1,1,22,22,16]
a: [3,27,0,15,1,23]
s: 27
i: [1,14,0,11,1,0,0,1] t4_xboole_1
s: 14
d: [2,29,1,15,1,1,44,27,16]
a: [3,34,1,15,1,29]
s: 34
d: [2,44,1,15,1,1,44,27,14]
d: [2,50,1,15,1,1,22,44,16]
a: [3,54,1,15,1,50]
s: 54
d: [2,37,1,9,1,0,44,14,15]
s: 37
d: [2,57,2,9,1,0,44,37,16]
s: 57
d: [2,38,1,11,1,0,44,14,16]
s: 38
d: [2,39,1,11,1,0,44,14,16]
s: 39
d: [2,91,2,15,1,1,44,27,38]
d: [2,108,2,15,1,1,22,91,14]
a: [3,112,2,15,1,108]
s: 112
d: [2,41,1,11,1,0,44,14,15]
s: 41
d: [2,58,2,9,1,0,44,37,16]
s: 58
d: [2,42,1,11,1,0,44,14,16]
s: 42
r: 34
r: 27
r: 112
r: 54
d: [2,122,2,15,1,1,23,112,42]
d: [2,196,2,15,1,1,22,122,42]
d: [2,197,2,0,0,1,19,196]
e: 197
f: 197,108,23,13
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t113_xboole_1
% SZS output start Proof for xboole_1__t113_xboole_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3) = k2_xboole_0(X0,k2_xboole_0(k2_xboole_0(X1,X2),X3)) [input t113_xboole_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3) = k2_xboole_0(X0,k2_xboole_0(k2_xboole_0(X1,X2),X3)) [negated conjecture 1]
3. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
6. ! [X0] : ! [X1] : ! [X2] : k2_xboole_0(k2_xboole_0(X0,X1),X2) = k2_xboole_0(X0,k2_xboole_0(X1,X2)) [input t4_xboole_1]
7. ~! [X0,X1,X2,X3] : k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3) = k2_xboole_0(X0,k2_xboole_0(k2_xboole_0(X1,X2),X3)) [flattening 2]
8. ! [X0,X1,X2] : k2_xboole_0(k2_xboole_0(X0,X1),X2) = k2_xboole_0(X0,k2_xboole_0(X1,X2)) [flattening 6]
10. ? [X0,X1,X2,X3] : k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3) != k2_xboole_0(X0,k2_xboole_0(k2_xboole_0(X1,X2),X3)) [ennf transformation 7]
11. ? [X0,X1,X2,X3] : k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3) != k2_xboole_0(X0,k2_xboole_0(k2_xboole_0(X1,X2),X3)) => k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) != k2_xboole_0(sK0,k2_xboole_0(k2_xboole_0(sK1,sK2),sK3)) [choice axiom]
12. k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) != k2_xboole_0(sK0,k2_xboole_0(k2_xboole_0(sK1,sK2),sK3)) [skolemisation 10,11]
13. k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) != k2_xboole_0(sK0,k2_xboole_0(k2_xboole_0(sK1,sK2),sK3)) [cnf transformation 12] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
14. k2_xboole_0(k2_xboole_0(X0,X1),X2) = k2_xboole_0(X0,k2_xboole_0(X1,X2)) [cnf transformation 8] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
16. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
18. 1 <=> k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) = k2_xboole_0(sK0,k2_xboole_0(k2_xboole_0(sK1,sK2),sK3)) [avatar definition]
20. k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) != k2_xboole_0(sK0,k2_xboole_0(k2_xboole_0(sK1,sK2),sK3)) <- (~1) [avatar component clause 18] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
21. ~1 [avatar split clause 13,18]
22. k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) != k2_xboole_0(sK0,k2_xboole_0(sK3,k2_xboole_0(sK1,sK2))) <- (~1) [backward demodulation 20,16] {a:0,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
23. k2_xboole_0(sK0,k2_xboole_0(sK3,k2_xboole_0(sK1,sK2))) != k2_xboole_0(k2_xboole_0(sK2,k2_xboole_0(sK0,sK1)),sK3) <- (~1) [forward demodulation 22,16] {a:0,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
25. 2 <=> k2_xboole_0(sK0,k2_xboole_0(sK3,k2_xboole_0(sK1,sK2))) = k2_xboole_0(k2_xboole_0(sK2,k2_xboole_0(sK0,sK1)),sK3) [avatar definition]
27. k2_xboole_0(sK0,k2_xboole_0(sK3,k2_xboole_0(sK1,sK2))) != k2_xboole_0(k2_xboole_0(sK2,k2_xboole_0(sK0,sK1)),sK3) <- (~2) [avatar component clause 25] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
28. ~2 | 1 [avatar split clause 23,18,25]
38. k2_xboole_0(X2,k2_xboole_0(X3,X4)) = k2_xboole_0(k2_xboole_0(X3,X2),X4) [superposition 14,16] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
42. k2_xboole_0(X5,k2_xboole_0(X6,X7)) = k2_xboole_0(X7,k2_xboole_0(X5,X6)) [superposition 14,16] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
91. k2_xboole_0(sK0,k2_xboole_0(sK3,k2_xboole_0(sK1,sK2))) != k2_xboole_0(k2_xboole_0(sK0,sK1),k2_xboole_0(sK2,sK3)) <- (~2) [superposition 27,38] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
108. k2_xboole_0(sK0,k2_xboole_0(sK3,k2_xboole_0(sK1,sK2))) != k2_xboole_0(sK0,k2_xboole_0(sK1,k2_xboole_0(sK2,sK3))) <- (~2) [forward demodulation 91,14] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
110. 5 <=> k2_xboole_0(sK0,k2_xboole_0(sK3,k2_xboole_0(sK1,sK2))) = k2_xboole_0(sK0,k2_xboole_0(sK1,k2_xboole_0(sK2,sK3))) [avatar definition]
112. k2_xboole_0(sK0,k2_xboole_0(sK3,k2_xboole_0(sK1,sK2))) != k2_xboole_0(sK0,k2_xboole_0(sK1,k2_xboole_0(sK2,sK3))) <- (~5) [avatar component clause 110] {a:2,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
113. ~5 | 2 [avatar split clause 108,25,110]
122. k2_xboole_0(sK0,k2_xboole_0(sK1,k2_xboole_0(sK2,sK3))) != k2_xboole_0(sK0,k2_xboole_0(sK2,k2_xboole_0(sK3,sK1))) <- (~5) [backward demodulation 112,42] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
196. k2_xboole_0(sK0,k2_xboole_0(sK1,k2_xboole_0(sK2,sK3))) != k2_xboole_0(sK0,k2_xboole_0(sK1,k2_xboole_0(sK2,sK3))) <- (~5) [forward demodulation 122,42] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
197. $false <- (~5) [trivial inequality removal 196] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
198. 5 [avatar contradiction clause 197]
224. $false [avatar sat refutation 21,28,113,198] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:41,thDist:-41}
% SZS output end Proof for xboole_1__t113_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 5
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 4
% Generated clauses: 179
% Active clauses: 12
% Passive clauses: 42
% Final active clauses: 7
% Final passive clauses: 27
% Main loop iterations started: 15
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw demodulations: 54
% Bw demodulations: 5
% 
% >>> Deletion Inferences
% Forward subsumptions: 33
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Forward superposition: 59
% Backward superposition: 41
% Self superposition: 10
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 2
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% forward demodulation: 0.001 s ( own 0.0000 s ) 
% other: 0.002 s ( own 0.001 s ) 

