i: [1,33,0,9,2,1,0,0] t25_gate_1
a: [3,75,0,2,1,33]
s: 75
i: [1,37,0,9,2,1,0,0] t25_gate_1
a: [3,56,0,7,1,37]
s: 56
i: [1,38,0,2,1,0,0,255] fc1_xboole_0
a: [3,82,0,2,1,38]
s: 82
i: [1,47,0,2,1,0,0,1] rc2_xboole_0
a: [3,87,0,2,1,47]
s: 87
i: [1,48,0,2,1,0,0,1] rc1_xboole_0
a: [3,92,0,2,1,48]
s: 92
i: [1,39,0,5,2,0,0,255] fc2_gate_1
s: 39
i: [1,40,0,5,2,0,0,255] fc1_gate_1
s: 40
i: [1,41,0,11,2,0,0,255] d19_gate_1
s: 41
r: 75
a: [3,51,0,2,1,37]
s: 51
a: [3,76,0,2,1,33]
s: 76
i: [1,42,0,11,2,0,0,255] d19_gate_1
s: 42
d: [2,95,1,5,2,1,44,56,41]
d: [2,96,1,3,1,2,21,95,75]
a: [3,100,1,3,1,96]
s: 100
d: [2,105,2,0,0,2,52,82,100,40]
e: 105
d: [2,106,2,2,1,1,40,100,40]
d: [2,108,2,0,0,2,21,106,82]
e: 108
r: 100
i: [1,34,0,9,2,1,0,0] t25_gate_1
a: [3,71,0,2,1,34]
s: 71
a: [3,99,1,3,1,96]
s: 99
i: [1,43,0,11,2,0,0,255] d19_gate_1
s: 43
i: [1,35,0,9,2,1,0,0] t25_gate_1
a: [3,66,0,2,1,35]
s: 66
i: [1,44,0,11,2,0,0,255] d19_gate_1
s: 44
i: [1,36,0,9,2,1,0,0] t25_gate_1
a: [3,61,0,2,1,36]
s: 61
i: [1,45,0,11,2,0,0,255] d19_gate_1
s: 45
d: [2,117,1,5,2,1,44,56,45]
d: [2,118,1,2,1,2,21,117,99]
d: [2,119,1,0,0,3,21,118,51]
e: 119
r: 51
r: 56
a: [3,52,0,2,1,37]
s: 52
a: [3,55,0,7,1,37]
s: 55
i: [1,46,0,18,6,0,0,255] d19_gate_1
s: 46
d: [2,136,1,12,6,1,44,55,46]
d: [2,137,1,10,5,2,21,136,82]
d: [2,138,1,8,4,3,21,137,52]
d: [2,139,1,6,3,4,21,138,61]
d: [2,140,1,4,2,5,21,139,66]
d: [2,141,1,2,1,6,21,140,71]
d: [2,142,1,0,0,7,21,141,76]
e: 142
i: [1,26,0,17,6,1,0,0] t25_gate_1
d: [2,103,1,5,2,1,44,56,42]
d: [2,111,1,5,2,1,44,56,43]
d: [2,112,1,2,1,2,21,111,99]
d: [2,114,1,5,2,1,44,56,44]
d: [2,115,1,2,1,2,21,114,99]
f: 142,119,115,112,108,103,95,38,26,33,34,35,36,37
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t25_gate_1
% SZS output start Proof for gate_1__t25_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : (~(v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)) & ~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & ~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)))) [input t25_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : (~(v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)) & ~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & ~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : (~(k19_gate_1(X0,X1,X2,X3,X4) != k1_xboole_0 & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) & (~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k19_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0))) [input d19_gate_1]
14. ~! [X0,X1,X2,X3,X4] : (~(v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)) & ~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & ~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)))) [flattening 2]
15. ! [X0,X1,X2,X3,X4] : (~(k19_gate_1(X0,X1,X2,X3,X4) != k1_xboole_0 & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) & (~(v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k19_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0))) [flattening 13]
16. ? [X0,X1,X2,X3,X4] : ((v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)) & (~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) | (v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2,X3,X4] : ((k19_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) & (k19_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0) | (v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)))) [ennf transformation 15]
20. ? [X0,X1,X2,X3,X4] : ((v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)) & (~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) | (v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k19_gate_1(X0,X1,X2,X3,X4)))) => ((v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) & (~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0))) | (v1_xboole_0(sK4) & v1_xboole_0(sK3) & v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0) & ~v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)))) [choice axiom]
21. (v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) & (~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0))) | (v1_xboole_0(sK4) & v1_xboole_0(sK3) & v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0) & ~v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4))) [skolemisation 16,20]
26. ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | ~v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) [cnf transformation 21] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) | v1_xboole_0(sK0) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) | v1_xboole_0(sK1) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) | v1_xboole_0(sK2) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) | v1_xboole_0(sK3) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) | v1_xboole_0(sK4) [cnf transformation 21] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
40. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
41. k19_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. k19_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X1) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. k19_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X2) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. k19_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X3) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
45. k19_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X4) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
46. k19_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:18,wCS:7200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
50. 1 <=> v1_xboole_0(sK4) [avatar definition]
51. ~v1_xboole_0(sK4) <- (~1) [avatar component clause 50] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
52. v1_xboole_0(sK4) <- (1) [avatar component clause 50] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
54. 2 <=> v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) [avatar definition]
55. ~v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) <- (~2) [avatar component clause 54] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. v1_xboole_0(k19_gate_1(sK0,sK1,sK2,sK3,sK4)) <- (2) [avatar component clause 54] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
57. 1 | 2 [avatar split clause 37,54,50]
59. 3 <=> v1_xboole_0(sK3) [avatar definition]
61. v1_xboole_0(sK3) <- (3) [avatar component clause 59] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. 3 | 2 [avatar split clause 36,54,59]
64. 4 <=> v1_xboole_0(sK2) [avatar definition]
66. v1_xboole_0(sK2) <- (4) [avatar component clause 64] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
67. 4 | 2 [avatar split clause 35,54,64]
69. 5 <=> v1_xboole_0(sK1) [avatar definition]
71. v1_xboole_0(sK1) <- (5) [avatar component clause 69] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. 5 | 2 [avatar split clause 34,54,69]
74. 6 <=> v1_xboole_0(sK0) [avatar definition]
76. v1_xboole_0(sK0) <- (6) [avatar component clause 74] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
77. 6 | 2 [avatar split clause 33,54,74]
78. ~2 | ~6 | ~5 | ~4 | ~3 | ~1 [avatar split clause 26,50,59,64,69,74,54]
80. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
82. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 80] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
83. 7 [avatar split clause 38,80]
95. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK0) <- (2) [superposition 56,41] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. 10 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
99. ~v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (~10) [avatar component clause 98] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (10) [avatar component clause 98] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
102. 6 | 10 | ~2 [avatar split clause 95,54,98,74]
103. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK1) <- (2) [superposition 56,42] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. 5 | 10 | ~2 [avatar split clause 103,54,98,69]
106. ~v1_xboole_0(k1_xboole_0) <- (10) [resolution 100,40] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
108. $false <- (7, 10) [subsumption resolution 106,82] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
109. ~7 | ~10 [avatar contradiction clause 108]
111. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK2) <- (2) [superposition 56,43] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
112. v1_xboole_0(sK2) <- (2, ~10) [subsumption resolution 111,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. 4 | ~2 | 10 [avatar split clause 112,98,54,64]
114. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK3) <- (2) [superposition 56,44] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
115. v1_xboole_0(sK3) <- (2, ~10) [subsumption resolution 114,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. 3 | ~2 | 10 [avatar split clause 115,98,54,59]
117. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK4) <- (2) [superposition 56,45] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. v1_xboole_0(sK4) <- (2, ~10) [subsumption resolution 117,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. $false <- (~1, 2, ~10) [subsumption resolution 118,51] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
120. 1 | ~2 | 10 [avatar contradiction clause 119]
136. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~2) [superposition 55,46] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
137. ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~2, 7) [subsumption resolution 136,82] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
138. ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 7) [subsumption resolution 137,52] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
139. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 3, 7) [subsumption resolution 138,61] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
140. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 3, 4, 7) [subsumption resolution 139,66] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
141. ~v1_xboole_0(sK0) <- (1, ~2, 3, 4, 5, 7) [subsumption resolution 140,71] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. $false <- (1, ~2, 3, 4, 5, 6, 7) [subsumption resolution 141,76] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
143. ~1 | 2 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar contradiction clause 142]
144. $false [avatar sat refutation 57,62,67,72,77,78,83,102,104,109,113,116,120,143] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:81,thDist:-81}
% SZS output end Proof for gate_1__t25_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 75
% Active clauses: 22
% Passive clauses: 22
% Final active clauses: 18
% Main loop iterations started: 22
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 1
% Forward superposition: 10
% Backward superposition: 11
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 20
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 3
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

