i: [1,35,0,11,2,1,0,0] t33_gate_1
a: [3,91,0,2,1,35]
s: 91
i: [1,42,0,2,1,0,0,255] fc1_xboole_0
a: [3,98,0,2,1,42]
s: 98
i: [1,53,0,2,1,0,0,1] rc2_xboole_0
a: [3,103,0,2,1,53]
s: 103
i: [1,54,0,2,1,0,0,1] rc1_xboole_0
a: [3,108,0,2,1,54]
s: 108
i: [1,41,0,11,2,1,0,0] t33_gate_1
a: [3,62,0,9,1,41]
s: 62
i: [1,43,0,5,2,0,0,255] fc2_gate_1
s: 43
i: [1,44,0,5,2,0,0,255] fc1_gate_1
s: 44
i: [1,45,0,13,2,0,0,255] d27_gate_1
s: 45
r: 91
a: [3,57,0,2,1,41]
s: 57
a: [3,92,0,2,1,35]
s: 92
i: [1,46,0,13,2,0,0,255] d27_gate_1
s: 46
i: [1,36,0,11,2,1,0,0] t33_gate_1
a: [3,87,0,2,1,36]
s: 87
i: [1,47,0,13,2,0,0,255] d27_gate_1
s: 47
d: [2,111,1,5,2,1,44,62,45]
d: [2,112,1,3,1,2,21,111,91]
a: [3,116,1,3,1,112]
s: 116
d: [2,123,2,0,0,2,52,98,116,44]
e: 123
d: [2,124,2,2,1,1,40,116,44]
d: [2,126,2,0,0,2,21,124,98]
e: 126
r: 116
i: [1,37,0,11,2,1,0,0] t33_gate_1
a: [3,82,0,2,1,37]
s: 82
i: [1,48,0,13,2,0,0,255] d27_gate_1
s: 48
i: [1,38,0,11,2,1,0,0] t33_gate_1
a: [3,77,0,2,1,38]
s: 77
i: [1,49,0,13,2,0,0,255] d27_gate_1
s: 49
i: [1,39,0,11,2,1,0,0] t33_gate_1
a: [3,72,0,2,1,39]
s: 72
i: [1,50,0,13,2,0,0,255] d27_gate_1
s: 50
i: [1,40,0,11,2,1,0,0] t33_gate_1
a: [3,67,0,2,1,40]
s: 67
a: [3,115,1,3,1,112]
s: 115
i: [1,51,0,13,2,0,0,255] d27_gate_1
s: 51
d: [2,135,1,5,2,1,44,62,51]
d: [2,136,1,2,1,2,21,135,115]
d: [2,137,1,0,0,3,21,136,57]
e: 137
r: 57
r: 62
a: [3,58,0,2,1,41]
s: 58
a: [3,61,0,9,1,41]
s: 61
i: [1,52,0,24,8,0,0,255] d27_gate_1
s: 52
d: [2,160,1,16,8,1,44,61,52]
d: [2,161,1,14,7,2,21,160,98]
d: [2,162,1,12,6,3,21,161,58]
d: [2,163,1,10,5,4,21,162,67]
d: [2,164,1,8,4,5,21,163,72]
d: [2,165,1,6,3,6,21,164,77]
d: [2,166,1,4,2,7,21,165,82]
d: [2,167,1,2,1,8,21,166,87]
d: [2,168,1,0,0,9,21,167,92]
e: 168
i: [1,26,0,23,8,1,0,0] t33_gate_1
d: [2,119,1,5,2,1,44,62,46]
d: [2,121,1,5,2,1,44,62,47]
d: [2,128,1,5,2,1,44,62,48]
d: [2,130,1,5,2,1,44,62,49]
d: [2,132,1,5,2,1,44,62,50]
f: 168,137,132,130,128,126,121,119,111,42,26,35,36,37,38,39,40,41
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t33_gate_1
% SZS output start Proof for gate_1__t33_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : (~(v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)) & ~(v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & ~(v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)))) [input t33_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : (~(v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)) & ~(v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & ~(v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : (~(k27_gate_1(X0,X1,X2,X3,X4,X5,X6) != k1_xboole_0 & v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) & (~(v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0))) [input d27_gate_1]
14. ~! [X0,X1,X2,X3,X4,X5,X6] : (~(v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)) & ~(v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & ~(v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)))) [flattening 2]
15. ! [X0,X1,X2,X3,X4,X5,X6] : (~(k27_gate_1(X0,X1,X2,X3,X4,X5,X6) != k1_xboole_0 & v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) & (~(v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0))) [flattening 13]
16. ? [X0,X1,X2,X3,X4,X5,X6] : ((v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)) & (~v1_xboole_0(X6) | ~v1_xboole_0(X5) | ~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) | (v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2,X3,X4,X5,X6] : ((k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_xboole_0 | ~v1_xboole_0(X6) | ~v1_xboole_0(X5) | ~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) & (k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0) | (v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)))) [ennf transformation 15]
20. ? [X0,X1,X2,X3,X4,X5,X6] : ((v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)) & (~v1_xboole_0(X6) | ~v1_xboole_0(X5) | ~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) | (v1_xboole_0(X6) & v1_xboole_0(X5) & v1_xboole_0(X4) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0) & ~v1_xboole_0(k27_gate_1(X0,X1,X2,X3,X4,X5,X6)))) => ((v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) & (~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0))) | (v1_xboole_0(sK6) & v1_xboole_0(sK5) & v1_xboole_0(sK4) & v1_xboole_0(sK3) & v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0) & ~v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)))) [choice axiom]
21. (v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) & (~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0))) | (v1_xboole_0(sK6) & v1_xboole_0(sK5) & v1_xboole_0(sK4) & v1_xboole_0(sK3) & v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0) & ~v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6))) [skolemisation 16,20]
26. ~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | ~v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) [cnf transformation 21] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) | v1_xboole_0(sK0) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) | v1_xboole_0(sK1) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) | v1_xboole_0(sK2) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) | v1_xboole_0(sK3) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) | v1_xboole_0(sK4) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) | v1_xboole_0(sK5) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
41. v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) | v1_xboole_0(sK6) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
44. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
45. k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
46. k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X1) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
47. k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X2) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
48. k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X3) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
49. k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X4) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
50. k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X5) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
51. k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X6) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
52. k27_gate_1(X0,X1,X2,X3,X4,X5,X6) = k1_xboole_0 | ~v1_xboole_0(X6) | ~v1_xboole_0(X5) | ~v1_xboole_0(X4) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:24,wCS:9600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
56. 1 <=> v1_xboole_0(sK6) [avatar definition]
57. ~v1_xboole_0(sK6) <- (~1) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
58. v1_xboole_0(sK6) <- (1) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 2 <=> v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) [avatar definition]
61. ~v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) <- (~2) [avatar component clause 60] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. v1_xboole_0(k27_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) <- (2) [avatar component clause 60] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. 1 | 2 [avatar split clause 41,60,56]
65. 3 <=> v1_xboole_0(sK5) [avatar definition]
67. v1_xboole_0(sK5) <- (3) [avatar component clause 65] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. 3 | 2 [avatar split clause 40,60,65]
70. 4 <=> v1_xboole_0(sK4) [avatar definition]
72. v1_xboole_0(sK4) <- (4) [avatar component clause 70] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
73. 4 | 2 [avatar split clause 39,60,70]
75. 5 <=> v1_xboole_0(sK3) [avatar definition]
77. v1_xboole_0(sK3) <- (5) [avatar component clause 75] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
78. 5 | 2 [avatar split clause 38,60,75]
80. 6 <=> v1_xboole_0(sK2) [avatar definition]
82. v1_xboole_0(sK2) <- (6) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. 6 | 2 [avatar split clause 37,60,80]
85. 7 <=> v1_xboole_0(sK1) [avatar definition]
87. v1_xboole_0(sK1) <- (7) [avatar component clause 85] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
88. 7 | 2 [avatar split clause 36,60,85]
90. 8 <=> v1_xboole_0(sK0) [avatar definition]
92. v1_xboole_0(sK0) <- (8) [avatar component clause 90] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
93. 8 | 2 [avatar split clause 35,60,90]
94. ~2 | ~8 | ~7 | ~6 | ~5 | ~4 | ~3 | ~1 [avatar split clause 26,56,65,70,75,80,85,90,60]
96. 9 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
98. v1_xboole_0(k1_xboole_0) <- (9) [avatar component clause 96] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
99. 9 [avatar split clause 42,96]
111. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK0) <- (2) [superposition 62,45] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 12 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
115. ~v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (~12) [avatar component clause 114] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (12) [avatar component clause 114] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. 8 | 12 | ~2 [avatar split clause 111,60,114,90]
119. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK1) <- (2) [superposition 62,46] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
120. 7 | 12 | ~2 [avatar split clause 119,60,114,85]
121. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK2) <- (2) [superposition 62,47] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
122. 6 | 12 | ~2 [avatar split clause 121,60,114,80]
124. ~v1_xboole_0(k1_xboole_0) <- (12) [resolution 116,44] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
126. $false <- (9, 12) [subsumption resolution 124,98] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
127. ~9 | ~12 [avatar contradiction clause 126]
128. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK3) <- (2) [superposition 62,48] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. 5 | 12 | ~2 [avatar split clause 128,60,114,75]
130. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK4) <- (2) [superposition 62,49] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
131. 4 | 12 | ~2 [avatar split clause 130,60,114,70]
132. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK5) <- (2) [superposition 62,50] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
133. 3 | 12 | ~2 [avatar split clause 132,60,114,65]
135. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK6) <- (2) [superposition 62,51] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
136. v1_xboole_0(sK6) <- (2, ~12) [subsumption resolution 135,115] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
137. $false <- (~1, 2, ~12) [subsumption resolution 136,57] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
138. 1 | ~2 | 12 [avatar contradiction clause 137]
160. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~2) [superposition 61,52] {a:1,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
161. ~v1_xboole_0(sK6) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~2, 9) [subsumption resolution 160,98] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
162. ~v1_xboole_0(sK5) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 9) [subsumption resolution 161,58] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
163. ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 3, 9) [subsumption resolution 162,67] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
164. ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 3, 4, 9) [subsumption resolution 163,72] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
165. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 3, 4, 5, 9) [subsumption resolution 164,77] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
166. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, ~2, 3, 4, 5, 6, 9) [subsumption resolution 165,82] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
167. ~v1_xboole_0(sK0) <- (1, ~2, 3, 4, 5, 6, 7, 9) [subsumption resolution 166,87] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
168. $false <- (1, ~2, 3, 4, 5, 6, 7, 8, 9) [subsumption resolution 167,92] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
169. ~1 | 2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 [avatar contradiction clause 168]
170. $false [avatar sat refutation 63,68,73,78,83,88,93,94,99,118,120,122,127,129,131,133,138,169] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:103,thDist:-103}
% SZS output end Proof for gate_1__t33_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 91
% Active clauses: 26
% Passive clauses: 26
% Final active clauses: 22
% Main loop iterations started: 26
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 23
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 1
% Forward superposition: 14
% Backward superposition: 15
% 
% >>> AVATAR
% Split clauses: 14
% Split components: 34
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 3
% SAT solver binary clauses: 9
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

