i: [1,38,0,6,1,1,0,0] t61_zfmisc_1
a: [3,64,0,6,1,38]
s: 64
i: [1,39,0,7,1,1,0,0] t61_zfmisc_1
a: [3,59,0,7,1,39]
s: 59
i: [1,40,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,69,0,2,1,40]
s: 69
i: [1,41,0,2,1,0,0,1] fc1_xboole_0
a: [3,74,0,2,1,41]
s: 74
i: [1,54,0,2,1,0,0,2] rc2_xboole_0
a: [3,84,0,2,1,54]
s: 84
i: [1,55,0,2,1,0,0,2] rc1_xboole_0
a: [3,89,0,2,1,55]
s: 89
i: [1,43,0,3,1,0,0,1] fc2_xboole_0
s: 43
i: [1,42,0,3,1,0,0,255] d2_xboole_0
a: [3,79,0,3,1,42]
s: 79
r: 69
i: [1,44,0,5,1,0,0,1] t4_boole
s: 44
i: [1,45,0,5,1,0,0,1] t3_boole
s: 45
i: [1,46,0,5,2,0,0,1] t6_boole
s: 46
i: [1,52,0,5,2,0,0,2] t7_boole
s: 52
i: [1,53,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 53
i: [1,51,0,7,3,0,0,1] t8_boole
s: 51
i: [1,47,0,9,2,0,0,1] l28_zfmisc_1
s: 47
i: [1,48,0,9,2,0,0,1] l28_zfmisc_1
s: 48
d: [2,98,1,7,2,1,44,59,48]
a: [3,106,1,3,1,98]
s: 106
i: [1,49,0,10,2,0,0,1] l26_zfmisc_1
s: 49
d: [2,114,1,8,2,0,44,49,45]
d: [2,115,1,3,1,0,19,114]
s: 115
i: [1,50,0,10,2,0,0,1] l26_zfmisc_1
s: 50
d: [2,118,2,0,0,2,52,106,59,50]
e: 118
d: [2,120,1,8,2,1,44,59,50]
d: [2,126,1,3,1,1,19,120]
d: [2,128,1,0,0,2,21,126,106]
e: 128
d: [2,99,1,6,2,1,44,64,48]
d: [2,102,1,3,1,1,19,99]
f: 128,102,38,39
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t61_zfmisc_1
% SZS output start Proof for zfmisc_1__t61_zfmisc_1
1. ! [X0] : ! [X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) | k4_xboole_0(k1_tarski(X0),X1) = k1_xboole_0) [input t61_zfmisc_1]
2. ~! [X0] : ! [X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) | k4_xboole_0(k1_tarski(X0),X1) = k1_xboole_0) [negated conjecture 1]
18. ! [X0] : ! [X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) <=> ~r2_hidden(X0,X1)) [input l26_zfmisc_1]
19. ! [X0] : ! [X1] : (k4_xboole_0(k1_tarski(X0),X1) = k1_xboole_0 <=> r2_hidden(X0,X1)) [input l28_zfmisc_1]
20. ~! [X0,X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) | k4_xboole_0(k1_tarski(X0),X1) = k1_xboole_0) [flattening 2]
21. ! [X0,X1] : (k4_xboole_0(k1_tarski(X0),X1) = k1_xboole_0 <=> r2_hidden(X0,X1)) [flattening 19]
22. ! [X0,X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) <=> ~r2_hidden(X0,X1)) [flattening 18]
25. ? [X0,X1] : (k1_tarski(X0) != k4_xboole_0(k1_tarski(X0),X1) & k4_xboole_0(k1_tarski(X0),X1) != k1_xboole_0) [ennf transformation 20]
30. ? [X0,X1] : (k1_tarski(X0) != k4_xboole_0(k1_tarski(X0),X1) & k4_xboole_0(k1_tarski(X0),X1) != k1_xboole_0) => (k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),sK1) & k1_xboole_0 != k4_xboole_0(k1_tarski(sK0),sK1)) [choice axiom]
31. k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),sK1) & k1_xboole_0 != k4_xboole_0(k1_tarski(sK0),sK1) [skolemisation 25,30]
32. ! [X0,X1] : ((k4_xboole_0(k1_tarski(X0),X1) = k1_xboole_0 | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X1) | k4_xboole_0(k1_tarski(X0),X1) != k1_xboole_0)) [nnf transformation 21]
33. ! [X0,X1] : ((k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1)) & (~r2_hidden(X0,X1) | k1_tarski(X0) != k4_xboole_0(k1_tarski(X0),X1))) [nnf transformation 22]
38. k1_xboole_0 != k4_xboole_0(k1_tarski(sK0),sK1) [cnf transformation 31] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),sK1) [cnf transformation 31] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. k4_xboole_0(k1_tarski(X0),X1) = k1_xboole_0 | ~r2_hidden(X0,X1) [cnf transformation 32] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1) [cnf transformation 33] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. 1 <=> k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),sK1) [avatar definition]
59. k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),sK1) <- (~1) [avatar component clause 57] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. ~1 [avatar split clause 39,57]
62. 2 <=> k1_xboole_0 = k4_xboole_0(k1_tarski(sK0),sK1) [avatar definition]
64. k1_xboole_0 != k4_xboole_0(k1_tarski(sK0),sK1) <- (~2) [avatar component clause 62] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. ~2 [avatar split clause 38,62]
99. k1_xboole_0 != k1_xboole_0 | ~r2_hidden(sK0,sK1) <- (~2) [superposition 64,48] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
102. ~r2_hidden(sK0,sK1) <- (~2) [trivial inequality removal 99] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
104. 8 <=> r2_hidden(sK0,sK1) [avatar definition]
106. ~r2_hidden(sK0,sK1) <- (~8) [avatar component clause 104] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
112. ~8 | 2 [avatar split clause 102,62,104]
120. k1_tarski(sK0) != k1_tarski(sK0) | r2_hidden(sK0,sK1) <- (~1) [superposition 59,50] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
126. r2_hidden(sK0,sK1) <- (~1) [trivial inequality removal 120] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
128. $false <- (~1, ~8) [subsumption resolution 126,106] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
129. 1 | 8 [avatar contradiction clause 128]
131. $false [avatar sat refutation 60,65,112,129] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:38,thDist:-38}
% SZS output end Proof for zfmisc_1__t61_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 18
% Generated clauses: 54
% Active clauses: 20
% Passive clauses: 27
% Final active clauses: 19
% Final passive clauses: 7
% Main loop iterations started: 20
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Trivial inequalities: 3
% Fw subsumption resolutions: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 2
% 
% >>> Generating Inferences
% Unit resulting resolution: 1
% Forward superposition: 6
% Backward superposition: 12
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 7
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

