i: [1,300,0,2,1,1,0,0] t143_xreal_1
a: [3,477,0,2,1,300]
s: 477
i: [1,299,0,2,1,1,0,0] t143_xreal_1
a: [3,482,0,2,1,299]
s: 482
i: [1,303,0,6,2,1,0,0] t143_xreal_1
a: [3,458,0,3,1,303]
s: 458
i: [1,302,0,6,2,1,0,0] t143_xreal_1
a: [3,463,0,3,1,302]
s: 463
i: [1,301,0,5,1,1,0,0] t143_xreal_1
a: [3,472,0,5,1,301]
s: 472
i: [1,304,0,2,1,0,0,255] fc1_numbers
a: [3,487,0,2,1,304]
s: 487
i: [1,305,0,2,1,0,0,255] spc1_boole
a: [3,492,0,2,1,305]
s: 492
i: [1,306,0,2,1,0,0,255] fc1_xboole_0
a: [3,497,0,2,1,306]
s: 497
i: [1,316,0,2,1,0,0,255] fc6_ordinal1
a: [3,542,0,2,1,316]
s: 542
i: [1,315,0,2,1,0,0,255] fc6_ordinal1
a: [3,547,0,2,1,315]
s: 547
i: [1,317,0,2,1,0,0,255] spc1_numerals
a: [3,567,0,2,1,317]
s: 567
i: [1,408,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,571,0,2,1,408]
s: 571
i: [1,418,0,2,1,0,0,2] rc2_xboole_0
a: [3,579,0,2,1,418]
s: 579
i: [1,419,0,2,1,0,0,255] rc4_ordinal1
a: [3,584,0,2,1,419]
s: 584
i: [1,420,0,2,1,0,0,255] rc1_ordinal1
a: [3,589,0,2,1,420]
s: 589
i: [1,421,0,2,1,0,0,1] rc1_xreal_0
a: [3,594,0,2,1,421]
s: 594
i: [1,422,0,2,1,0,0,2] rc1_xxreal_0
a: [3,599,0,2,1,422]
s: 599
d: [2,741,0,0,0,2,21,599,571]
e: 741
r: 571
s: 599
i: [1,423,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,604,0,2,1,423]
s: 604
i: [1,424,0,2,1,0,0,2] rc1_xboole_0
a: [3,609,0,2,1,424]
s: 609
i: [1,428,0,2,1,0,0,255] rc2_ordinal1
a: [3,614,0,2,1,428]
s: 614
i: [1,427,0,2,1,0,0,255] rc2_ordinal1
a: [3,619,0,2,1,427]
s: 619
i: [1,426,0,2,1,0,0,255] rc2_ordinal1
a: [3,624,0,2,1,426]
s: 624
i: [1,425,0,2,1,0,0,255] rc2_ordinal1
a: [3,629,0,2,1,425]
s: 629
i: [1,430,0,2,1,0,0,255] rc5_ordinal1
a: [3,634,0,2,1,430]
s: 634
i: [1,429,0,2,1,0,0,255] rc5_ordinal1
a: [3,639,0,2,1,429]
s: 639
i: [1,432,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,644,0,2,1,432]
s: 644
i: [1,431,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,649,0,2,1,431]
s: 649
i: [1,434,0,2,1,0,0,2] rc3_xxreal_0
a: [3,654,0,2,1,434]
s: 654
i: [1,433,0,2,1,0,0,2] rc3_xxreal_0
a: [3,659,0,2,1,433]
s: 659
i: [1,436,0,2,1,0,0,2] rc2_xxreal_0
a: [3,664,0,2,1,436]
s: 664
i: [1,435,0,2,1,0,0,2] rc2_xxreal_0
a: [3,669,0,2,1,435]
s: 669
i: [1,440,0,2,1,0,0,2] rc3_xreal_0
a: [3,674,0,2,1,440]
s: 674
i: [1,439,0,2,1,0,0,2] rc3_xreal_0
a: [3,679,0,2,1,439]
s: 679
i: [1,438,0,2,1,0,0,2] rc3_xreal_0
a: [3,684,0,2,1,438]
s: 684
i: [1,437,0,2,1,0,0,2] rc3_xreal_0
a: [3,689,0,2,1,437]
s: 689
i: [1,444,0,2,1,0,0,2] rc2_xreal_0
a: [3,694,0,2,1,444]
s: 694
i: [1,443,0,2,1,0,0,2] rc2_xreal_0
a: [3,699,0,2,1,443]
s: 699
i: [1,442,0,2,1,0,0,2] rc2_xreal_0
a: [3,704,0,2,1,442]
s: 704
i: [1,441,0,2,1,0,0,2] rc2_xreal_0
a: [3,709,0,2,1,441]
s: 709
i: [1,446,0,2,1,0,0,2] rc4_xxreal_0
a: [3,714,0,2,1,446]
s: 714
i: [1,445,0,2,1,0,0,2] rc4_xxreal_0
a: [3,719,0,2,1,445]
s: 719
i: [1,450,0,2,1,0,0,2] rc4_xreal_0
a: [3,724,0,2,1,450]
s: 724
i: [1,449,0,2,1,0,0,2] rc4_xreal_0
a: [3,729,0,2,1,449]
s: 729
i: [1,448,0,2,1,0,0,2] rc4_xreal_0
a: [3,734,0,2,1,448]
s: 734
i: [1,447,0,2,1,0,0,2] rc4_xreal_0
a: [3,739,0,2,1,447]
s: 739
i: [1,390,0,3,1,0,0,255] reflexivity_r1_tarski
s: 390
i: [1,307,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,502,0,3,1,307]
s: 502
i: [1,308,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,507,0,3,1,308]
s: 507
r: 497
d: [2,743,0,2,1,2,23,497,507]
a: [3,747,0,2,1,743]
s: 747
i: [1,309,0,3,1,0,0,255] t1_numerals
a: [3,512,0,3,1,309]
s: 512
i: [1,310,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,517,0,3,1,310]
s: 517
i: [1,320,0,3,1,0,0,255] spc1_numerals
a: [3,552,0,3,1,320]
s: 552
i: [1,319,0,3,1,0,0,255] spc1_numerals
a: [3,557,0,3,1,319]
s: 557
d: [2,749,0,3,1,2,22,512,507]
a: [3,753,0,3,1,749]
s: 753
d: [2,755,0,3,1,2,22,557,502]
a: [3,759,0,3,1,755]
s: 759
i: [1,322,0,4,2,0,0,255] cc6_ordinal1
s: 322
i: [1,323,0,4,2,0,0,255] cc2_xreal_0
s: 323
d: [2,761,1,2,1,1,40,322,584]
a: [3,766,1,2,1,761]
s: 766
i: [1,324,0,4,2,0,0,255] cc2_xxreal_0
s: 324
i: [1,325,0,4,2,0,0,255] cc2_xcmplx_0
s: 325
d: [2,762,1,2,1,1,40,322,634]
a: [3,771,1,2,1,762]
s: 771
i: [1,326,0,4,2,0,0,255] cc1_ordinal1
s: 326
d: [2,773,1,2,1,1,40,323,584]
a: [3,778,1,2,1,773]
s: 778
i: [1,327,0,4,2,0,0,255] cc1_ordinal1
s: 327
i: [1,329,0,4,2,0,0,2] cc4_xreal_0
s: 329
d: [2,859,1,2,1,1,40,329,482]
a: [3,869,1,2,1,859]
s: 869
i: [1,330,0,4,2,0,0,1] cc3_xreal_0
s: 330
d: [2,860,1,2,1,1,40,329,477]
a: [3,874,1,2,1,860]
s: 874
i: [1,349,0,4,2,0,0,255] cc9_ordinal1
s: 349
i: [1,350,0,4,2,0,0,255] cc7_ordinal1
s: 350
d: [2,882,1,2,1,1,40,330,482]
a: [3,892,1,2,1,882]
s: 892
i: [1,351,0,4,2,0,0,255] cc3_ordinal1
s: 351
d: [2,883,1,2,1,1,40,330,477]
a: [3,897,1,2,1,883]
s: 897
i: [1,389,0,4,1,0,0,255] existence_m1_subset_1
s: 389
i: [1,311,0,4,1,0,0,255] dt_k5_numbers
a: [3,522,0,4,1,311]
s: 522
d: [2,774,1,2,1,1,40,323,634]
a: [3,783,1,2,1,774]
s: 783
i: [1,312,0,4,1,0,0,255] dt_k6_numbers
a: [3,527,0,4,1,312]
s: 527
d: [2,785,1,2,1,1,40,324,584]
a: [3,790,1,2,1,785]
s: 790
i: [1,318,0,4,1,0,0,255] spc1_numerals
a: [3,562,0,4,1,318]
s: 562
d: [2,977,0,4,1,2,22,522,502]
a: [3,981,0,4,1,977]
s: 981
d: [2,786,1,2,1,1,40,324,634]
a: [3,795,1,2,1,786]
s: 795
d: [2,987,0,4,1,2,22,527,502]
a: [3,991,0,4,1,987]
s: 991
d: [2,797,1,2,1,1,40,325,584]
a: [3,802,1,2,1,797]
s: 802
d: [2,993,0,4,1,2,22,562,502]
a: [3,997,0,4,1,993]
s: 997
i: [1,331,0,5,2,0,0,255] fc4_xreal_0
s: 331
d: [2,798,1,2,1,1,40,325,634]
a: [3,807,1,2,1,798]
s: 807
i: [1,332,0,5,2,0,0,255] fc4_xreal_0
s: 332
d: [2,809,1,2,1,1,40,326,542]
a: [3,817,1,2,1,809]
s: 817
i: [1,341,0,5,2,0,0,255] dt_k5_xcmplx_0
s: 341
i: [1,352,0,5,2,0,0,255] t6_boole
s: 352
d: [2,811,1,2,1,1,40,326,589]
a: [3,827,1,2,1,811]
s: 827
i: [1,353,0,5,2,0,0,255] cc8_ordinal1
s: 353
d: [2,834,1,2,1,1,40,327,542]
a: [3,842,1,2,1,834]
s: 842
i: [1,354,0,5,2,0,0,255] cc1_xreal_0
s: 354
i: [1,355,0,5,2,0,0,255] cc1_xcmplx_0
s: 355
d: [2,836,1,2,1,1,40,327,589]
a: [3,852,1,2,1,836]
s: 852
i: [1,386,0,5,2,0,0,255] t7_boole
s: 386
d: [2,862,1,2,1,1,40,329,594]
a: [3,880,1,2,1,862]
s: 880
i: [1,313,0,5,1,0,0,255] rqRealDiv__k7_xcmplx_0__r1_r1_r1
a: [3,532,0,5,1,313]
s: 532
i: [1,314,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,537,0,5,1,314]
s: 537
d: [2,885,1,2,1,1,40,330,594]
a: [3,903,1,2,1,885]
s: 903
a: [3,574,0,5,2,408]
s: 574
d: [2,905,1,2,1,1,40,349,747]
a: [3,912,1,2,1,905]
s: 912
d: [2,1001,0,5,2,1,22,352,507]
s: 1001
i: [1,366,0,6,3,0,0,255] cc2_ordinal1
s: 366
d: [2,1041,1,5,2,2,44,458,1001]
s: 1041
i: [1,369,0,6,3,0,0,2] cc5_xxreal_0
s: 369
d: [2,1042,1,5,2,2,44,463,1001]
s: 1042
i: [1,371,0,6,3,0,0,2] cc5_xxreal_0
s: 371
i: [1,372,0,6,3,0,0,2] cc3_xxreal_0
s: 372
d: [2,1043,1,7,2,2,44,472,1001]
s: 1043
i: [1,374,0,6,3,0,0,2] cc3_xxreal_0
s: 374
d: [2,1094,2,8,2,3,40,1043,574]
a: [3,1098,2,4,1,1094]
s: 1098
i: [1,378,0,6,3,0,0,2] cc7_xxreal_0
s: 378
i: [1,379,0,6,3,0,0,2] cc7_xxreal_0
s: 379
d: [2,906,1,2,1,1,40,349,609]
a: [3,917,1,2,1,906]
s: 917
i: [1,382,0,6,2,0,0,255] t1_subset
s: 382
d: [2,907,1,2,1,1,40,349,719]
a: [3,922,1,2,1,907]
s: 922
i: [1,391,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 391
i: [1,328,0,7,3,0,0,255] cc5_ordinal1
s: 328
d: [2,908,1,2,1,1,40,349,739]
a: [3,927,1,2,1,908]
s: 927
i: [1,342,0,7,2,0,0,1] t5_arithm
s: 342
d: [2,929,1,2,1,1,40,350,747]
a: [3,936,1,2,1,929]
s: 936
i: [1,343,0,7,2,0,0,1] t2_arithm
s: 343
i: [1,344,0,7,2,0,0,255] involutiveness_k5_xcmplx_0
s: 344
d: [2,930,1,2,1,1,40,350,609]
a: [3,941,1,2,1,930]
s: 941
i: [1,345,0,7,2,0,0,255] t6_arithm
s: 345
d: [2,931,1,2,1,1,40,350,719]
a: [3,946,1,2,1,931]
s: 946
i: [1,346,0,7,2,0,0,255] t3_arithm
s: 346
i: [1,356,0,7,3,0,0,255] fc30_xreal_0
s: 356
d: [2,932,1,2,1,1,40,350,739]
a: [3,951,1,2,1,932]
s: 951
i: [1,357,0,7,3,0,0,255] fc30_xreal_0
s: 357
d: [2,953,1,2,1,1,40,351,747]
a: [3,960,1,2,1,953]
s: 960
i: [1,362,0,7,3,0,0,255] fc7_xcmplx_0
s: 362
i: [1,363,0,7,3,0,0,255] fc7_xcmplx_0
s: 363
d: [2,954,1,2,1,1,40,351,609]
a: [3,965,1,2,1,954]
s: 965
i: [1,364,0,7,3,0,0,255] fc28_xreal_0
s: 364
d: [2,955,1,2,1,1,40,351,719]
a: [3,970,1,2,1,955]
s: 970
i: [1,365,0,7,3,0,0,255] fc28_xreal_0
s: 365
i: [1,375,0,7,3,0,0,255] fc29_xreal_0
s: 375
d: [2,956,1,2,1,1,40,351,739]
a: [3,975,1,2,1,956]
s: 975
i: [1,376,0,7,3,0,0,255] fc29_xreal_0
s: 376
d: [2,1003,1,2,1,1,40,353,759]
a: [3,1009,1,2,1,1003]
s: 1009
i: [1,380,0,7,3,0,0,255] fc27_xreal_0
s: 380
i: [1,381,0,7,3,0,0,255] fc27_xreal_0
s: 381
d: [2,1016,1,2,1,1,40,354,552]
a: [3,1021,1,2,1,1016]
s: 1021
i: [1,384,0,7,2,0,0,255] t3_subset
s: 384
d: [2,1028,1,2,1,1,40,355,552]
a: [3,1033,1,2,1,1028]
s: 1033
i: [1,385,0,7,3,0,0,2] t8_boole
s: 385
i: [1,321,0,8,3,0,0,255] cc10_ordinal1
s: 321
d: [2,1051,1,4,2,1,40,369,654]
d: [2,1053,1,2,1,2,21,1051,659]
a: [3,1057,1,2,1,1053]
s: 1057
i: [1,347,0,8,2,0,0,255] spc3_arithm
s: 347
d: [2,1052,1,4,2,1,40,369,679]
d: [2,1059,1,2,1,2,21,1052,684]
a: [3,1063,1,2,1,1059]
s: 1063
i: [1,359,0,8,4,0,0,2] cc4_xxreal_0
s: 359
i: [1,361,0,8,4,0,0,2] cc6_xxreal_0
s: 361
d: [2,1297,2,6,3,1,40,359,869]
a: [3,1312,2,2,1,1297]
s: 1312
r: 1312
i: [1,367,0,8,4,0,0,2] cc8_xxreal_0
s: 367
a: [3,1311,2,2,1,1297]
s: 1311
i: [1,383,0,8,3,0,0,255] t2_subset
s: 383
i: [1,406,0,8,3,0,0,1] fc8_xreal_0
s: 406
a: [3,1316,2,2,1,1297]
s: 1316
i: [1,407,0,8,3,0,0,2] fc6_xreal_0
s: 407
d: [2,1298,2,6,3,1,40,359,874]
a: [3,1325,2,2,1,1298]
s: 1325
r: 1325
i: [1,410,0,8,3,0,0,1] fc5_xcmplx_0
s: 410
i: [1,411,0,8,3,0,0,2] fc3_xcmplx_0
s: 411
a: [3,1324,2,2,1,1298]
s: 1324
i: [1,388,0,9,3,0,0,255] t5_subset
s: 388
a: [3,1329,2,2,1,1298]
s: 1329
i: [1,387,0,10,3,0,0,255] t4_subset
s: 387
i: [1,409,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 409
d: [2,1065,1,4,2,1,40,371,654]
d: [2,1067,1,2,1,2,21,1065,659]
a: [3,1071,1,2,1,1067]
s: 1071
i: [1,335,0,11,5,0,0,1] t2_real
s: 335
a: [3,1319,2,2,1,1297]
s: 1319
i: [1,336,0,11,5,0,0,1] t1_real
s: 336
i: [1,337,0,11,5,0,0,1] t8_real
s: 337
a: [3,1332,2,2,1,1298]
s: 1332
i: [1,338,0,11,5,0,0,1] t7_real
s: 338
d: [2,1529,1,12,4,1,40,337,472]
a: [3,1538,1,4,1,1529]
s: 1538
d: [2,1552,2,0,0,3,52,482,477,1538,406]
e: 1552
d: [2,1553,2,4,2,1,40,1538,406]
d: [2,1555,2,2,1,2,21,1553,477]
d: [2,1556,2,0,0,3,21,1555,482]
e: 1556
r: 1538
i: [1,339,0,11,5,0,0,1] t3_real
s: 339
i: [1,340,0,11,5,0,0,1] t4_real
s: 340
a: [3,1537,1,4,1,1529]
s: 1537
d: [2,1558,3,0,0,2,52,1098,1537,329]
e: 1558
d: [2,1560,2,4,1,1,40,1537,329]
d: [2,1562,2,0,0,2,21,1560,1098]
e: 1562
r: 1098
i: [1,412,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 412
a: [3,1546,1,4,1,1529]
s: 1546
i: [1,348,0,12,3,0,0,255] d9_xcmplx_0
s: 348
i: [1,392,0,12,5,0,0,1] fc33_xreal_0
s: 392
d: [2,1640,3,0,0,5,52,1319,482,1332,477,1546,392]
e: 1640
d: [2,1641,2,8,4,1,40,392,1546]
d: [2,1648,2,6,3,2,21,1641,477]
d: [2,1649,2,4,2,3,21,1648,1332]
d: [2,1650,2,2,1,4,21,1649,482]
d: [2,1651,2,0,0,5,21,1650,1319]
e: 1651
r: 1319
d: [2,1492,1,8,4,1,40,335,458]
d: [2,1501,1,6,3,2,21,1492,482]
a: [3,1505,1,2,1,1501]
s: 1505
r: 1042
r: 463
r: 1332
r: 1505
a: [3,462,0,3,1,302]
s: 462
a: [3,1504,1,2,1,1501]
s: 1504
a: [3,454,0,3,1,303]
s: 454
r: 1329
a: [3,467,0,3,1,302]
s: 467
r: 1311
r: 1316
a: [3,1541,1,2,1,1529]
s: 1541
i: [1,393,0,12,5,0,0,2] fc26_xreal_0
s: 393
s: 1312
i: [1,394,0,12,5,0,0,1] fc31_xreal_0
s: 394
i: [1,395,0,12,5,0,0,1] fc32_xreal_0
s: 395
a: [3,1315,2,2,1,1297]
s: 1315
i: [1,400,0,12,5,0,0,1] fc9_xcmplx_0
s: 400
s: 1319
i: [1,401,0,12,5,0,0,2] fc8_xcmplx_0
s: 401
i: [1,402,0,12,5,0,0,2] fc23_xreal_0
s: 402
a: [3,1328,2,2,1,1298]
s: 1328
i: [1,403,0,12,5,0,0,2] fc24_xreal_0
s: 403
a: [3,1333,2,2,1,1298]
s: 1333
i: [1,404,0,12,5,0,0,1] fc34_xreal_0
s: 404
d: [2,1761,3,0,0,5,52,1315,482,1328,477,1546,404]
e: 1761
d: [2,1762,2,8,4,1,40,404,1546]
d: [2,1769,2,6,3,2,21,1762,477]
d: [2,1770,2,4,2,3,21,1769,1328]
d: [2,1771,2,2,1,4,21,1770,482]
d: [2,1772,2,0,0,5,21,1771,1315]
e: 1772
r: 1315
r: 1328
r: 1541
i: [1,405,0,12,5,0,0,2] fc25_xreal_0
s: 405
a: [3,1542,1,2,1,1529]
s: 1542
r: 1041
r: 458
r: 462
r: 467
r: 1319
r: 1324
r: 1333
r: 1542
a: [3,457,0,3,1,303]
s: 457
d: [2,1002,1,2,1,1,40,353,753]
d: [2,1124,2,2,1,1,40,936,324]
d: [2,1125,2,2,1,1,40,936,323]
d: [2,1493,1,8,4,1,40,335,463]
d: [2,1507,1,6,3,2,21,1493,477]
d: [2,1655,2,4,2,1,40,1505,369]
d: [2,1658,2,2,1,2,21,1655,747]
d: [2,1660,2,4,2,3,21,1648,482]
d: [2,1670,1,8,4,1,40,454,336]
d: [2,1676,1,6,3,2,21,1670,477]
d: [2,1678,1,8,4,1,40,467,336]
d: [2,1684,1,6,3,2,21,1678,482]
d: [2,1774,2,4,2,3,21,1769,482]
d: [2,1759,3,4,2,1,40,1333,371]
d: [2,1776,3,2,1,2,21,1759,874]
d: [2,1795,2,4,2,1,40,1542,372]
d: [2,1797,2,2,1,2,21,1795,747]
d: [2,1800,1,7,3,1,40,457,409]
d: [2,1804,1,5,2,2,21,1800,869]
f: 1804,1797,1776,1774,1684,1676,1660,1658,1556,1529,1507,1501,1125,1124,1002,860,859,749,743,309,308,306,299,300,301,302,303
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t143_xreal_1
% SZS output start Proof for xreal_1__t143_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~(~r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(X1,k6_numbers)) & ~(~r1_xxreal_0(X0,k6_numbers) & ~r1_xxreal_0(k6_numbers,X1)) & ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1))))) [input t143_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~(~r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(X1,k6_numbers)) & ~(~r1_xxreal_0(X0,k6_numbers) & ~r1_xxreal_0(k6_numbers,X1)) & ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1))))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
14. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
19. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
28. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
44. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
55. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc33_xreal_0]
56. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc34_xreal_0]
65. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
85. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
100. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
101. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
104. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
120. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
132. ? [X0] : (? [X1] : (((r1_xxreal_0(k6_numbers,X0) | r1_xxreal_0(X1,k6_numbers)) & (r1_xxreal_0(X0,k6_numbers) | r1_xxreal_0(k6_numbers,X1)) & ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
133. ? [X0] : (? [X1] : ((r1_xxreal_0(k6_numbers,X0) | r1_xxreal_0(X1,k6_numbers)) & (r1_xxreal_0(X0,k6_numbers) | r1_xxreal_0(k6_numbers,X1)) & ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 132]
136. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
137. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
141. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
148. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 104]
149. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 148]
150. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 101]
151. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 150]
152. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 120]
153. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 152]
172. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
189. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 19]
190. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 189]
191. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 14]
192. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 191]
209. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 55]
210. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 209]
231. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 56]
232. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 231]
235. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 65]
236. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 235]
241. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 28]
242. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 241]
259. ? [X0] : (? [X1] : ((r1_xxreal_0(k6_numbers,X0) | r1_xxreal_0(X1,k6_numbers)) & (r1_xxreal_0(X0,k6_numbers) | r1_xxreal_0(k6_numbers,X1)) & ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : ((r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(X1,k6_numbers)) & (r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(k6_numbers,X1)) & ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,X1)) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
260. ? [X1] : ((r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(X1,k6_numbers)) & (r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(k6_numbers,X1)) & ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,X1)) & v1_xreal_0(X1)) => ((r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(sK1,k6_numbers)) & (r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(k6_numbers,sK1)) & ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) & v1_xreal_0(sK1)) [choice axiom]
261. ((r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(sK1,k6_numbers)) & (r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(k6_numbers,sK1)) & ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 133,260,259]
299. v1_xreal_0(sK0) [cnf transformation 261] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
300. v1_xreal_0(sK1) [cnf transformation 261] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
301. ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) [cnf transformation 261] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
302. r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(k6_numbers,sK1) [cnf transformation 261] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
303. r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(sK1,k6_numbers) [cnf transformation 261] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
306. v1_xboole_0(k1_xboole_0) [cnf transformation 44] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
308. k6_numbers = k1_xboole_0 [cnf transformation 85] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
309. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 100] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
323. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 136] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
324. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 137] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
329. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 141] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
335. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 149] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
336. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 151] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
337. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 153] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
353. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 172] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
369. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 190] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
371. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 190] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
372. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 192] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
392. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 210] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
404. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 232] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
406. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 236] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
409. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 242] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
452. 1 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
454. r1_xxreal_0(sK1,k6_numbers) <- (1) [avatar component clause 452] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
456. 2 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
457. ~r1_xxreal_0(k6_numbers,sK0) <- (~2) [avatar component clause 456] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
458. r1_xxreal_0(k6_numbers,sK0) <- (2) [avatar component clause 456] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
459. 1 | 2 [avatar split clause 303,456,452]
461. 3 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
463. r1_xxreal_0(k6_numbers,sK1) <- (3) [avatar component clause 461] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
465. 4 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
467. r1_xxreal_0(sK0,k6_numbers) <- (4) [avatar component clause 465] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
468. 3 | 4 [avatar split clause 302,465,461]
470. 5 <=> r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) [avatar definition]
472. ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) <- (~5) [avatar component clause 470] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
473. ~5 [avatar split clause 301,470]
475. 6 <=> v1_xreal_0(sK1) [avatar definition]
477. v1_xreal_0(sK1) <- (6) [avatar component clause 475] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
478. 6 [avatar split clause 300,475]
480. 7 <=> v1_xreal_0(sK0) [avatar definition]
482. v1_xreal_0(sK0) <- (7) [avatar component clause 480] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
483. 7 [avatar split clause 299,480]
495. 10 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
497. v1_xboole_0(k1_xboole_0) <- (10) [avatar component clause 495] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
498. 10 [avatar split clause 306,495]
505. 12 <=> k6_numbers = k1_xboole_0 [avatar definition]
507. k6_numbers = k1_xboole_0 <- (12) [avatar component clause 505] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
508. 12 [avatar split clause 308,505]
510. 13 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
512. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (13) [avatar component clause 510] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
513. 13 [avatar split clause 309,510]
743. v1_xboole_0(k6_numbers) <- (10, 12) [backward demodulation 497,507] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
745. 60 <=> v1_xboole_0(k6_numbers) [avatar definition]
747. v1_xboole_0(k6_numbers) <- (60) [avatar component clause 745] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
748. 60 | ~10 | ~12 [avatar split clause 743,505,495,745]
749. m1_subset_1(k6_numbers,k4_ordinal1) <- (12, 13) [forward demodulation 512,507] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
751. 61 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
753. m1_subset_1(k6_numbers,k4_ordinal1) <- (61) [avatar component clause 751] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
754. 61 | ~12 | ~13 [avatar split clause 749,510,505,751]
859. v1_xxreal_0(sK0) <- (7) [resolution 329,482] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
860. v1_xxreal_0(sK1) <- (6) [resolution 329,477] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
867. 79 <=> v1_xxreal_0(sK0) [avatar definition]
869. v1_xxreal_0(sK0) <- (79) [avatar component clause 867] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
870. 79 | ~7 [avatar split clause 859,480,867]
872. 80 <=> v1_xxreal_0(sK1) [avatar definition]
874. v1_xxreal_0(sK1) <- (80) [avatar component clause 872] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
875. 80 | ~6 [avatar split clause 860,475,872]
934. 89 <=> v7_ordinal1(k6_numbers) [avatar definition]
936. v7_ordinal1(k6_numbers) <- (89) [avatar component clause 934] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1002. v7_ordinal1(k6_numbers) <- (61) [resolution 353,753] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1005. 89 | ~61 [avatar split clause 1002,751,934]
1124. v1_xxreal_0(k6_numbers) <- (89) [resolution 936,324] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1125. v1_xreal_0(k6_numbers) <- (89) [resolution 936,323] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1134. 117 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1137. 117 | ~89 [avatar split clause 1124,934,1134]
1139. 118 <=> v1_xreal_0(k6_numbers) [avatar definition]
1142. 118 | ~89 [avatar split clause 1125,934,1139]
1314. 137 <=> v2_xxreal_0(sK0) [avatar definition]
1318. 138 <=> v3_xxreal_0(sK0) [avatar definition]
1327. 140 <=> v2_xxreal_0(sK1) [avatar definition]
1331. 141 <=> v3_xxreal_0(sK1) [avatar definition]
1333. v3_xxreal_0(sK1) <- (141) [avatar component clause 1331] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1492. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (2) [resolution 335,458] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1493. ~v3_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (3) [resolution 335,463] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1501. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 7) [subsumption resolution 1492,482] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1503. 164 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1505. v3_xxreal_0(k6_numbers) <- (164) [avatar component clause 1503] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1506. ~118 | 164 | ~138 | ~2 | ~7 [avatar split clause 1501,480,456,1318,1503,1139]
1507. ~v3_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 6) [subsumption resolution 1493,477] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1508. ~118 | 164 | ~141 | ~3 | ~6 [avatar split clause 1507,475,461,1331,1503,1139]
1529. v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (~5) [resolution 337,472] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1536. 166 <=> v1_xreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1538. ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~166) [avatar component clause 1536] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1540. 167 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1542. v2_xxreal_0(k6_numbers) <- (167) [avatar component clause 1540] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1544. 168 <=> v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1546. v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) <- (168) [avatar component clause 1544] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1547. ~118 | ~166 | 167 | 168 | 5 [avatar split clause 1529,470,1544,1540,1536,1139]
1553. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~166) [resolution 1538,406] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1555. ~v1_xreal_0(sK0) <- (6, ~166) [subsumption resolution 1553,477] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1556. $false <- (6, 7, ~166) [subsumption resolution 1555,482] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1557. ~6 | ~7 | 166 [avatar contradiction clause 1556]
1641. ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (168) [resolution 392,1546] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1648. v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (6, 168) [subsumption resolution 1641,477] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1655. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (164) [resolution 1505,369] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1658. ~v1_xxreal_0(k6_numbers) <- (60, 164) [subsumption resolution 1655,747] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1659. ~117 | ~60 | ~164 [avatar split clause 1658,1503,745,1134]
1660. v3_xxreal_0(sK1) | v3_xxreal_0(sK0) <- (6, 7, 168) [subsumption resolution 1648,482] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1661. 138 | 141 | ~6 | ~7 | ~168 [avatar split clause 1660,1544,480,475,1331,1318]
1670. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (1) [resolution 454,336] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1676. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (1, 6) [subsumption resolution 1670,477] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1677. ~118 | 167 | ~140 | ~1 | ~6 [avatar split clause 1676,475,452,1327,1540,1139]
1678. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (4) [resolution 467,336] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1684. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (4, 7) [subsumption resolution 1678,482] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1685. ~118 | 167 | ~137 | ~4 | ~7 [avatar split clause 1684,480,465,1314,1540,1139]
1759. ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (141) [resolution 1333,371] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1762. ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (168) [resolution 404,1546] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1769. v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (6, 168) [subsumption resolution 1762,477] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1774. v2_xxreal_0(sK1) | v2_xxreal_0(sK0) <- (6, 7, 168) [subsumption resolution 1769,482] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1775. 137 | 140 | ~6 | ~7 | ~168 [avatar split clause 1774,1544,480,475,1327,1314]
1776. ~v2_xxreal_0(sK1) <- (80, 141) [subsumption resolution 1759,874] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1777. ~140 | ~80 | ~141 [avatar split clause 1776,1331,872,1327]
1795. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (167) [resolution 1542,372] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1797. ~v1_xxreal_0(k6_numbers) <- (60, 167) [subsumption resolution 1795,747] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1798. ~117 | ~60 | ~167 [avatar split clause 1797,1540,745,1134]
1800. r1_xxreal_0(sK0,k6_numbers) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(sK0) <- (~2) [resolution 457,409] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1804. r1_xxreal_0(sK0,k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (~2, 79) [subsumption resolution 1800,869] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1805. ~117 | 4 | 2 | ~79 [avatar split clause 1804,867,456,465,1134]
1806. $false [avatar sat refutation 459,468,473,478,483,498,508,513,748,754,870,875,1005,1137,1142,1506,1508,1547,1557,1659,1661,1677,1685,1775,1777,1798,1805] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1145,thDist:-1145}
% SZS output end Proof for xreal_1__t143_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 119
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 152
% Generated clauses: 802
% Active clauses: 188
% Passive clauses: 339
% Final active clauses: 163
% Final passive clauses: 125
% Main loop iterations started: 208
% 
% >>> Simplifying Inferences
% Duplicate literals: 65
% Fw subsumption resolutions: 77
% Fw demodulations: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 49
% Forward subsumptions: 86
% 
% >>> Generating Inferences
% Binary resolution: 189
% Unit resulting resolution: 4
% Factoring: 1
% Forward superposition: 91
% Backward superposition: 25
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 42
% Split components: 102
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 213
% SAT solver unit clauses: 56
% SAT solver binary clauses: 91
% 
% Memory used [KB]: 5884
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.008 s ( own 0.005 s ) 

