i: [1,146,0,2,1,1,0,0] l20_real_lat
a: [3,225,0,2,1,146]
s: 225
i: [1,145,0,2,1,1,0,0] l20_real_lat
a: [3,230,0,2,1,145]
s: 230
i: [1,147,0,3,1,1,0,0] l20_real_lat
a: [3,220,0,3,1,147]
s: 220
i: [1,150,0,4,1,1,0,0] l20_real_lat
a: [3,205,0,4,1,150]
s: 205
i: [1,148,0,4,1,1,0,0] l20_real_lat
a: [3,215,0,4,1,148]
s: 215
i: [1,149,0,6,1,1,0,0] l20_real_lat
a: [3,210,0,6,1,149]
s: 210
i: [1,151,0,2,1,0,0,2] fc1_xboole_0
a: [3,235,0,2,1,151]
s: 235
i: [1,190,0,2,1,0,0,1] rc2_xboole_0
a: [3,240,0,2,1,190]
s: 240
i: [1,191,0,2,1,0,0,1] rc1_xboole_0
a: [3,245,0,2,1,191]
s: 245
i: [1,193,0,2,1,0,0,255] rc2_membered
a: [3,250,0,2,1,193]
s: 250
i: [1,192,0,2,1,0,0,255] rc2_membered
a: [3,255,0,2,1,192]
s: 255
i: [1,195,0,2,1,0,0,255] rc1_membered
a: [3,260,0,2,1,195]
s: 260
i: [1,194,0,2,1,0,0,255] rc1_membered
a: [3,265,0,2,1,194]
s: 265
i: [1,197,0,2,1,0,0,255] rc3_membered
a: [3,270,0,2,1,197]
s: 270
i: [1,196,0,2,1,0,0,255] rc3_membered
a: [3,275,0,2,1,196]
s: 275
i: [1,198,0,2,1,0,0,255] rc4_xreal_0
a: [3,280,0,2,1,198]
s: 280
i: [1,177,0,3,1,0,0,255] reflexivity_r1_tarski
s: 177
i: [1,154,0,4,2,0,0,255] cc2_membered
s: 154
i: [1,155,0,4,2,0,0,255] cc3_membered
s: 155
i: [1,158,0,4,2,0,0,255] cc4_membered
s: 158
i: [1,159,0,4,2,0,0,255] cc5_membered
s: 159
i: [1,161,0,4,2,0,0,255] cc1_membered
s: 161
i: [1,163,0,4,2,0,0,255] cc18_membered
s: 163
d: [2,282,1,2,1,1,40,161,250]
a: [3,288,1,2,1,282]
s: 288
i: [1,176,0,4,1,0,0,1] existence_m1_subset_1
s: 176
i: [1,186,0,4,1,0,0,255] rc1_relset_1
s: 186
d: [2,283,1,2,1,1,40,161,260]
a: [3,293,1,2,1,283]
s: 293
i: [1,187,0,4,1,0,0,255] rc1_relset_1
s: 187
d: [2,284,1,2,1,1,40,161,270]
a: [3,298,1,2,1,284]
s: 298
i: [1,164,0,5,2,0,0,2] t6_boole
s: 164
i: [1,169,0,5,2,0,0,1] t7_boole
s: 169
d: [2,300,1,2,1,1,40,163,235]
a: [3,306,1,2,1,300]
s: 306
i: [1,188,0,5,1,0,0,255] rc1_relset_1
s: 188
d: [2,301,1,2,1,1,40,163,245]
a: [3,311,1,2,1,301]
s: 311
i: [1,165,0,6,2,0,0,1] t1_subset
s: 165
i: [1,180,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 180
d: [2,302,1,2,1,1,40,163,280]
a: [3,316,1,2,1,302]
s: 316
i: [1,167,0,7,2,0,0,255] t3_subset
s: 167
d: [2,318,2,2,1,1,40,288,154]
a: [3,322,2,2,1,318]
s: 322
i: [1,168,0,7,3,0,0,2] t8_boole
s: 168
i: [1,152,0,8,3,0,0,255] cc20_membered
s: 152
d: [2,365,1,6,2,1,40,152,210]
a: [3,370,1,4,1,365]
s: 370
i: [1,153,0,8,3,0,0,255] cc19_membered
s: 153
d: [2,377,1,6,2,1,40,153,210]
a: [3,382,1,4,1,377]
s: 382
i: [1,156,0,8,3,0,0,255] cc23_membered
s: 156
i: [1,157,0,8,3,0,0,255] cc22_membered
s: 157
d: [2,389,1,6,2,1,40,156,210]
a: [3,394,1,4,1,389]
s: 394
i: [1,160,0,8,3,0,0,255] cc21_membered
s: 160
d: [2,401,1,6,2,1,40,157,210]
a: [3,406,1,4,1,401]
s: 406
i: [1,162,0,8,3,0,0,255] cc24_membered
s: 162
i: [1,166,0,8,3,0,0,1] t2_subset
s: 166
d: [2,437,1,5,2,1,40,166,220]
d: [2,440,1,3,1,2,21,437,230]
a: [3,444,1,3,1,440]
s: 444
i: [1,183,0,8,2,0,0,2] cc1_relset_1
s: 183
d: [2,464,1,2,1,1,40,183,210]
a: [3,469,1,2,1,464]
s: 469
i: [1,185,0,8,1,0,0,255] rc1_relset_1
s: 185
i: [1,175,0,9,3,0,0,1] t5_subset
s: 175
d: [2,455,2,3,1,1,40,444,180]
a: [3,461,2,3,1,455]
s: 461
i: [1,184,0,9,2,0,0,255] cc2_relset_1
s: 184
d: [2,492,1,3,1,1,40,184,210]
a: [3,498,1,3,1,492]
s: 498
i: [1,174,0,10,3,0,0,1] t4_subset
s: 174
i: [1,178,0,10,3,0,0,1] cc4_relset_1
s: 178
d: [2,413,1,6,2,1,40,160,210]
a: [3,418,1,4,1,413]
s: 418
i: [1,179,0,10,3,0,0,1] cc3_relset_1
s: 179
d: [2,425,1,6,2,1,40,162,210]
a: [3,430,1,4,1,425]
s: 430
i: [1,200,0,10,2,0,0,1] d1_funct_2
s: 200
i: [1,181,0,11,3,0,0,1] redefinition_k1_relset_1
s: 181
d: [2,480,1,7,2,1,40,175,210]
a: [3,489,1,4,1,480]
s: 489
i: [1,182,0,11,3,0,0,2] dt_k1_relset_1
s: 182
d: [2,438,1,11,2,1,40,166,210]
a: [3,449,1,6,1,438]
s: 449
i: [1,189,0,11,2,0,0,1] fc8_relset_1
s: 189
i: [1,201,0,13,3,0,0,1] d1_funct_2
s: 201
a: [3,452,1,5,1,438]
s: 452
i: [1,170,0,18,4,0,0,1] d1_funct_2
s: 170
d: [2,549,1,12,3,1,40,170,210]
d: [2,552,1,8,2,2,21,549,215]
a: [3,556,1,5,1,552]
s: 556
d: [2,564,2,9,3,1,44,181,556]
d: [2,579,2,6,2,2,21,564,498]
d: [2,580,2,4,1,3,21,579,469]
f: 580,552,492,464,440,151,145,146,147,148,149,150
% Refutation found. Thanks to Tanya!
% SZS status Theorem for real_lat__l20_real_lat
% SZS output start Proof for real_lat__l20_real_lat
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,X0) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1) & v1_funct_1(X3)) => r2_hidden(X2,k9_xtuple_0(X3)))))) [input l20_real_lat]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,X0) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1) & v1_funct_1(X3)) => r2_hidden(X2,k9_xtuple_0(X3)))))) [negated conjecture 1]
13. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [input cc1_relset_1]
21. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [input cc2_relset_1]
30. ! [X0] : ! [X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ((k1_xboole_0 = X1 => (v1_funct_2(X2,X0,X1) <=> k1_xboole_0 = X2)) & (k1_xboole_0 != X1 => (v1_funct_2(X2,X0,X1) <=> k1_relset_1(X0,X2) = X0)))) [input d1_funct_2]
38. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
48. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => k1_relset_1(X0,X1) = k9_xtuple_0(X1)) [input redefinition_k1_relset_1]
51. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
59. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 51]
63. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ((k1_xboole_0 = X1 => (v1_funct_2(X2,X0,X1) <=> k1_xboole_0 = X2)) & (k1_xboole_0 != X1 => (v1_funct_2(X2,X0,X1) <=> k1_relset_1(X0,X2) = X0)))) [flattening 30]
67. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [flattening 13]
68. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [flattening 21]
70. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v4_relat_1(X2,X0)) [pure predicate removal 68]
87. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,X0) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1)) => r2_hidden(X2,k9_xtuple_0(X3)))))) [pure predicate removal 2]
88. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & (m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1))) & m1_subset_1(X2,X0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [ennf transformation 87]
89. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1)) & m1_subset_1(X2,X0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [flattening 88]
104. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 59]
105. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 104]
109. ! [X0,X1,X2] : ((((v1_funct_2(X2,X0,X1) <=> k1_xboole_0 = X2) | k1_xboole_0 != X1) & ((v1_funct_2(X2,X0,X1) <=> k1_relset_1(X0,X2) = X0) | k1_xboole_0 = X1)) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 63]
116. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 48]
117. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 116]
120. ! [X0,X1,X2] : (v1_relat_1(X2) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 67]
121. ! [X0,X1,X2] : (v4_relat_1(X2,X0) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 70]
123. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1)) & m1_subset_1(X2,X0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,X1))) & v1_funct_2(X3,sK0,X1)) & m1_subset_1(X2,sK0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
124. ? [X1] : (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,X1))) & v1_funct_2(X3,sK0,X1)) & m1_subset_1(X2,sK0)) & ~v1_xboole_0(X1)) => (? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(X3,sK0,sK1)) & m1_subset_1(X2,sK0)) & ~v1_xboole_0(sK1)) [choice axiom]
125. ? [X2] : (? [X3] : (~r2_hidden(X2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(X3,sK0,sK1)) & m1_subset_1(X2,sK0)) => (? [X3] : (~r2_hidden(sK2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(X3,sK0,sK1)) & m1_subset_1(sK2,sK0)) [choice axiom]
126. ? [X3] : (~r2_hidden(sK2,k9_xtuple_0(X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(X3,sK0,sK1)) => (~r2_hidden(sK2,k9_xtuple_0(sK3)) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(sK3,sK0,sK1)) [choice axiom]
127. (((~r2_hidden(sK2,k9_xtuple_0(sK3)) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(sK3,sK0,sK1)) & m1_subset_1(sK2,sK0)) & ~v1_xboole_0(sK1)) & ~v1_xboole_0(sK0) [skolemisation 89,126,125,124,123]
128. ! [X0,X1,X2] : (((((v1_funct_2(X2,X0,X1) | k1_xboole_0 != X2) & (k1_xboole_0 = X2 | ~v1_funct_2(X2,X0,X1))) | k1_xboole_0 != X1) & (((v1_funct_2(X2,X0,X1) | k1_relset_1(X0,X2) != X0) & (k1_relset_1(X0,X2) = X0 | ~v1_funct_2(X2,X0,X1))) | k1_xboole_0 = X1)) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [nnf transformation 109]
145. ~v1_xboole_0(sK0) [cnf transformation 127] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
146. ~v1_xboole_0(sK1) [cnf transformation 127] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
147. m1_subset_1(sK2,sK0) [cnf transformation 127] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
148. v1_funct_2(sK3,sK0,sK1) [cnf transformation 127] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
149. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [cnf transformation 127] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
150. ~r2_hidden(sK2,k9_xtuple_0(sK3)) [cnf transformation 127] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
151. v1_xboole_0(k1_xboole_0) [cnf transformation 38] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
166. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 105] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
170. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~v1_funct_2(X2,X0,X1) | k1_xboole_0 = X1 | k1_relset_1(X0,X2) = X0 [cnf transformation 128] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
181. k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 117] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
183. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v1_relat_1(X2) [cnf transformation 120] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
184. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v4_relat_1(X2,X0) [cnf transformation 121] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
203. 1 <=> r2_hidden(sK2,k9_xtuple_0(sK3)) [avatar definition]
206. ~1 [avatar split clause 150,203]
208. 2 <=> m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [avatar definition]
210. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) <- (2) [avatar component clause 208] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
211. 2 [avatar split clause 149,208]
213. 3 <=> v1_funct_2(sK3,sK0,sK1) [avatar definition]
215. v1_funct_2(sK3,sK0,sK1) <- (3) [avatar component clause 213] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
216. 3 [avatar split clause 148,213]
218. 4 <=> m1_subset_1(sK2,sK0) [avatar definition]
220. m1_subset_1(sK2,sK0) <- (4) [avatar component clause 218] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
221. 4 [avatar split clause 147,218]
223. 5 <=> v1_xboole_0(sK1) [avatar definition]
226. ~5 [avatar split clause 146,223]
228. 6 <=> v1_xboole_0(sK0) [avatar definition]
230. ~v1_xboole_0(sK0) <- (~6) [avatar component clause 228] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
231. ~6 [avatar split clause 145,228]
233. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
236. 7 [avatar split clause 151,233]
437. v1_xboole_0(sK0) | r2_hidden(sK2,sK0) <- (4) [resolution 166,220] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
440. r2_hidden(sK2,sK0) <- (4, ~6) [subsumption resolution 437,230] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
442. 42 <=> r2_hidden(sK2,sK0) [avatar definition]
445. 42 | ~4 | 6 [avatar split clause 440,228,218,442]
464. v1_relat_1(sK3) <- (2) [resolution 183,210] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
467. 46 <=> v1_relat_1(sK3) [avatar definition]
469. v1_relat_1(sK3) <- (46) [avatar component clause 467] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
470. 46 | ~2 [avatar split clause 464,208,467]
492. v4_relat_1(sK3,sK0) <- (2) [resolution 184,210] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
496. 49 <=> v4_relat_1(sK3,sK0) [avatar definition]
498. v4_relat_1(sK3,sK0) <- (49) [avatar component clause 496] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
499. 49 | ~2 [avatar split clause 492,208,496]
549. ~v1_funct_2(sK3,sK0,sK1) | k1_xboole_0 = sK1 | sK0 = k1_relset_1(sK0,sK3) <- (2) [resolution 170,210] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
552. k1_xboole_0 = sK1 | sK0 = k1_relset_1(sK0,sK3) <- (2, 3) [subsumption resolution 549,215] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
554. 51 <=> sK0 = k1_relset_1(sK0,sK3) [avatar definition]
556. sK0 = k1_relset_1(sK0,sK3) <- (51) [avatar component clause 554] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
558. 52 <=> k1_xboole_0 = sK1 [avatar definition]
561. 51 | 52 | ~2 | ~3 [avatar split clause 552,213,208,558,554]
564. sK0 = k9_xtuple_0(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (51) [superposition 181,556] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
568. 53 <=> sK0 = k9_xtuple_0(sK3) [avatar definition]
579. sK0 = k9_xtuple_0(sK3) | ~v1_relat_1(sK3) <- (49, 51) [subsumption resolution 564,498] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
580. sK0 = k9_xtuple_0(sK3) <- (46, 49, 51) [subsumption resolution 579,469] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
581. 53 | ~46 | ~49 | ~51 [avatar split clause 580,554,496,467,568]
582. k1_xboole_0 != sK1 | v1_xboole_0(sK1) | ~v1_xboole_0(k1_xboole_0) [theory tautology sat conflict] {a:0,w:7,wCS:2800,sine:255,thAx:0,allAx:1,thDist:-1}
583. sK0 != k9_xtuple_0(sK3) | r2_hidden(sK2,k9_xtuple_0(sK3)) | ~r2_hidden(sK2,sK0) [theory tautology sat conflict] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
584. $false [avatar sat refutation 206,211,216,221,226,231,236,445,470,499,561,581,582,583] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:305,thDist:-305}
% SZS output end Proof for real_lat__l20_real_lat
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 56
% >>> Preprocessing
% Introduced skolems: 12
% Pure predicates: 9
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 54
% Generated clauses: 216
% Active clauses: 74
% Passive clauses: 149
% Final active clauses: 74
% Final passive clauses: 75
% Main loop iterations started: 74
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 8
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 100
% Forward superposition: 5
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 18
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 51
% SAT solver unit clauses: 16
% SAT solver binary clauses: 18
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

