i: [1,72,0,8,2,1,0,0] t62_mcart_1
a: [3,107,0,4,1,72]
s: 107
i: [1,71,0,5,1,1,0,0] t62_mcart_1
a: [3,112,0,5,1,71]
s: 112
i: [1,70,0,5,1,1,0,0] t62_mcart_1
a: [3,117,0,5,1,70]
s: 117
i: [1,73,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,122,0,2,1,73]
s: 122
i: [1,74,0,2,1,0,0,1] fc1_xboole_0
a: [3,127,0,2,1,74]
s: 127
i: [1,93,0,2,1,0,0,2] rc2_xboole_0
a: [3,137,0,2,1,93]
s: 137
i: [1,94,0,2,1,0,0,2] rc1_xboole_0
a: [3,142,0,2,1,94]
s: 142
i: [1,96,0,2,1,0,0,2] rc1_relat_1
a: [3,147,0,2,1,96]
s: 147
i: [1,95,0,2,1,0,0,2] rc1_relat_1
a: [3,152,0,2,1,95]
s: 152
i: [1,97,0,2,1,0,0,255] rc2_relat_1
a: [3,157,0,2,1,97]
s: 157
i: [1,75,0,3,1,0,0,255] d2_xboole_0
a: [3,132,0,3,1,75]
s: 132
r: 122
i: [1,76,0,4,2,0,0,2] cc1_relat_1
s: 76
i: [1,89,0,4,1,0,0,1] existence_m1_subset_1
s: 89
d: [2,160,1,2,1,1,40,76,142]
a: [3,165,1,2,1,160]
s: 165
i: [1,90,0,4,1,0,0,1] fc6_relat_1
s: 90
i: [1,77,0,5,2,0,0,1] t6_boole
s: 77
d: [2,161,1,2,1,1,40,76,127]
a: [3,170,1,2,1,161]
s: 170
i: [1,88,0,5,2,0,0,255] t7_boole
s: 88
i: [1,98,0,5,1,0,0,1] t90_zfmisc_1
s: 98
i: [1,99,0,5,1,0,0,1] t90_zfmisc_1
s: 99
i: [1,80,0,6,2,0,0,255] t1_subset
s: 80
i: [1,91,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 91
i: [1,85,0,7,3,0,0,1] t8_boole
s: 85
d: [2,172,1,7,3,0,44,77,77]
s: 172
i: [1,81,0,8,3,0,0,255] t2_subset
s: 81
i: [1,92,0,8,3,0,0,1] fc10_subset_1
s: 92
d: [2,179,1,9,2,1,40,81,112]
a: [3,184,1,5,1,179]
s: 184
i: [1,82,0,11,3,0,0,1] t90_zfmisc_1
s: 82
a: [3,187,1,4,1,179]
s: 187
i: [1,86,0,15,4,0,0,1] t24_mcart_1
s: 86
i: [1,87,0,15,4,0,0,1] t24_mcart_1
s: 87
d: [2,192,2,5,1,1,40,184,91]
a: [3,198,2,5,1,192]
s: 198
d: [2,203,1,14,4,1,44,87,107]
d: [2,204,1,11,3,1,19,203]
s: 204
r: 204
r: 107
a: [3,103,0,4,1,72]
s: 103
a: [3,106,0,4,1,72]
s: 106
d: [2,217,1,14,4,1,44,86,103]
d: [2,218,1,11,3,1,19,217]
s: 218
d: [2,205,2,6,2,2,40,204,112]
a: [3,211,2,3,1,205]
s: 211
r: 198
r: 187
r: 184
r: 117
r: 112
d: [2,224,1,4,1,2,23,187,211]
d: [2,234,1,2,1,2,22,224,99]
d: [2,235,1,0,0,3,21,234,127]
e: 235
d: [2,226,0,5,1,2,23,117,211]
d: [2,238,0,0,0,2,21,226,99]
e: 238
r: 211
s: 112
a: [3,210,2,3,1,205]
s: 210
s: 117
s: 187
a: [3,215,2,3,1,205]
s: 215
r: 187
r: 117
r: 112
d: [2,245,1,4,1,2,23,187,215]
d: [2,248,1,2,1,2,22,245,98]
d: [2,249,1,0,0,3,21,248,127]
e: 249
d: [2,246,0,5,1,2,23,117,215]
d: [2,251,0,0,0,2,21,246,98]
e: 251
d: [2,241,2,6,2,2,40,112,218]
f: 251,241,238,205,70,71,72
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t62_mcart_1
% SZS output start Proof for mcart_1__t62_mcart_1
1. ! [X0] : ! [X1] : (k2_zfmisc_1(X0,X1) != k1_xboole_0 => ! [X2] : (m1_subset_1(X2,k2_zfmisc_1(X0,X1)) => (k2_xtuple_0(X2) != X2 & k1_xtuple_0(X2) != X2))) [input t62_mcart_1]
2. ~! [X0] : ! [X1] : (k2_zfmisc_1(X0,X1) != k1_xboole_0 => ! [X2] : (m1_subset_1(X2,k2_zfmisc_1(X0,X1)) => (k2_xtuple_0(X2) != X2 & k1_xtuple_0(X2) != X2))) [negated conjecture 1]
23. ! [X0] : ! [X1] : ~(~! [X2] : (m1_subset_1(X2,k2_zfmisc_1(X0,X1)) => (k2_xtuple_0(X2) != X2 & k1_xtuple_0(X2) != X2)) & k1_xboole_0 != X1 & k1_xboole_0 != X0) [input t24_mcart_1]
28. ! [X0] : ! [X1] : (k2_zfmisc_1(X0,X1) = k1_xboole_0 <=> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t90_zfmisc_1]
29. ~! [X0,X1] : (k2_zfmisc_1(X0,X1) != k1_xboole_0 => ! [X2] : (m1_subset_1(X2,k2_zfmisc_1(X0,X1)) => (k2_xtuple_0(X2) != X2 & k1_xtuple_0(X2) != X2))) [flattening 2]
32. ! [X0,X1] : (k2_zfmisc_1(X0,X1) = k1_xboole_0 <=> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 28]
34. ! [X0,X1] : ~(~! [X2] : (m1_subset_1(X2,k2_zfmisc_1(X0,X1)) => (k2_xtuple_0(X2) != X2 & k1_xtuple_0(X2) != X2)) & k1_xboole_0 != X1 & k1_xboole_0 != X0) [flattening 23]
39. ? [X0,X1] : (? [X2] : ((k2_xtuple_0(X2) = X2 | k1_xtuple_0(X2) = X2) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & k2_zfmisc_1(X0,X1) != k1_xboole_0) [ennf transformation 29]
50. ! [X0,X1] : (! [X2] : ((k2_xtuple_0(X2) != X2 & k1_xtuple_0(X2) != X2) | ~m1_subset_1(X2,k2_zfmisc_1(X0,X1))) | k1_xboole_0 = X1 | k1_xboole_0 = X0) [ennf transformation 34]
55. ? [X0,X1] : (? [X2] : ((k2_xtuple_0(X2) = X2 | k1_xtuple_0(X2) = X2) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & k2_zfmisc_1(X0,X1) != k1_xboole_0) => (? [X2] : ((k2_xtuple_0(X2) = X2 | k1_xtuple_0(X2) = X2) & m1_subset_1(X2,k2_zfmisc_1(sK0,sK1))) & k1_xboole_0 != k2_zfmisc_1(sK0,sK1)) [choice axiom]
56. ? [X2] : ((k2_xtuple_0(X2) = X2 | k1_xtuple_0(X2) = X2) & m1_subset_1(X2,k2_zfmisc_1(sK0,sK1))) => ((sK2 = k2_xtuple_0(sK2) | sK2 = k1_xtuple_0(sK2)) & m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1))) [choice axiom]
57. ((sK2 = k2_xtuple_0(sK2) | sK2 = k1_xtuple_0(sK2)) & m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1))) & k1_xboole_0 != k2_zfmisc_1(sK0,sK1) [skolemisation 39,56,55]
58. ! [X0,X1] : ((k2_zfmisc_1(X0,X1) = k1_xboole_0 | (k1_xboole_0 != X1 & k1_xboole_0 != X0)) & ((k1_xboole_0 = X1 | k1_xboole_0 = X0) | k2_zfmisc_1(X0,X1) != k1_xboole_0)) [nnf transformation 32]
59. ! [X0,X1] : ((k2_zfmisc_1(X0,X1) = k1_xboole_0 | (k1_xboole_0 != X1 & k1_xboole_0 != X0)) & (k1_xboole_0 = X1 | k1_xboole_0 = X0 | k2_zfmisc_1(X0,X1) != k1_xboole_0)) [flattening 58]
70. k1_xboole_0 != k2_zfmisc_1(sK0,sK1) [cnf transformation 57] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
71. m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1)) [cnf transformation 57] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
72. sK2 = k2_xtuple_0(sK2) | sK2 = k1_xtuple_0(sK2) [cnf transformation 57] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
83. k2_zfmisc_1(X0,X1) = k1_xboole_0 | k1_xboole_0 != X0 [cnf transformation 59] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
84. k2_zfmisc_1(X0,X1) = k1_xboole_0 | k1_xboole_0 != X1 [cnf transformation 59] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
86. k1_xtuple_0(X2) != X2 | ~m1_subset_1(X2,k2_zfmisc_1(X0,X1)) | k1_xboole_0 = X1 | k1_xboole_0 = X0 [cnf transformation 50] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. k2_xtuple_0(X2) != X2 | ~m1_subset_1(X2,k2_zfmisc_1(X0,X1)) | k1_xboole_0 = X1 | k1_xboole_0 = X0 [cnf transformation 50] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
98. k1_xboole_0 = k2_zfmisc_1(X0,k1_xboole_0) [equality resolution 84] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
99. k1_xboole_0 = k2_zfmisc_1(k1_xboole_0,X1) [equality resolution 83] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. 1 <=> sK2 = k1_xtuple_0(sK2) [avatar definition]
103. sK2 = k1_xtuple_0(sK2) <- (1) [avatar component clause 101] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 2 <=> sK2 = k2_xtuple_0(sK2) [avatar definition]
107. sK2 = k2_xtuple_0(sK2) <- (2) [avatar component clause 105] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
108. 1 | 2 [avatar split clause 72,105,101]
110. 3 <=> m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1)) [avatar definition]
112. m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1)) <- (3) [avatar component clause 110] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. 3 [avatar split clause 71,110]
115. 4 <=> k1_xboole_0 = k2_zfmisc_1(sK0,sK1) [avatar definition]
117. k1_xboole_0 != k2_zfmisc_1(sK0,sK1) <- (~4) [avatar component clause 115] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. ~4 [avatar split clause 70,115]
203. sK2 != sK2 | ~m1_subset_1(sK2,k2_zfmisc_1(X0,X1)) | k1_xboole_0 = X1 | k1_xboole_0 = X0 <- (2) [superposition 87,107] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
204. ~m1_subset_1(sK2,k2_zfmisc_1(X0,X1)) | k1_xboole_0 = X1 | k1_xboole_0 = X0 <- (2) [trivial inequality removal 203] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
205. k1_xboole_0 = sK1 | k1_xboole_0 = sK0 <- (2, 3) [resolution 204,112] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
209. 18 <=> k1_xboole_0 = sK0 [avatar definition]
211. k1_xboole_0 = sK0 <- (18) [avatar component clause 209] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
213. 19 <=> k1_xboole_0 = sK1 [avatar definition]
215. k1_xboole_0 = sK1 <- (19) [avatar component clause 213] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
216. 18 | 19 | ~2 | ~3 [avatar split clause 205,110,105,213,209]
217. sK2 != sK2 | ~m1_subset_1(sK2,k2_zfmisc_1(X0,X1)) | k1_xboole_0 = X1 | k1_xboole_0 = X0 <- (1) [superposition 86,103] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
218. ~m1_subset_1(sK2,k2_zfmisc_1(X0,X1)) | k1_xboole_0 = X1 | k1_xboole_0 = X0 <- (1) [trivial inequality removal 217] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
226. k1_xboole_0 != k2_zfmisc_1(k1_xboole_0,sK1) <- (~4, 18) [backward demodulation 117,211] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
238. $false <- (~4, 18) [subsumption resolution 226,99] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
239. 4 | ~18 [avatar contradiction clause 238]
241. k1_xboole_0 = sK1 | k1_xboole_0 = sK0 <- (1, 3) [resolution 112,218] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
243. 18 | 19 | ~1 | ~3 [avatar split clause 241,110,101,213,209]
246. k1_xboole_0 != k2_zfmisc_1(sK0,k1_xboole_0) <- (~4, 19) [backward demodulation 117,215] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
251. $false <- (~4, 19) [subsumption resolution 246,98] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
252. 4 | ~19 [avatar contradiction clause 251]
254. $false [avatar sat refutation 108,113,118,216,239,243,252] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:126,thDist:-126}
% SZS output end Proof for mcart_1__t62_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 27
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 28
% Generated clauses: 97
% Active clauses: 41
% Passive clauses: 47
% Final active clauses: 29
% Final passive clauses: 5
% Main loop iterations started: 42
% 
% >>> Simplifying Inferences
% Trivial inequalities: 2
% Fw subsumption resolutions: 4
% Fw demodulations: 6
% Bw demodulations: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 11
% Forward superposition: 6
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 27
% SAT solver unit clauses: 10
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

