i: [1,96,0,2,1,1,0,0] l3_polyred
a: [3,170,0,2,1,96]
s: 170
i: [1,95,0,2,1,1,0,0] l3_polyred
a: [3,175,0,2,1,95]
s: 175
i: [1,93,0,2,1,1,0,0] l3_polyred
a: [3,185,0,2,1,93]
s: 185
i: [1,91,0,2,1,1,0,0] l3_polyred
a: [3,195,0,2,1,91]
s: 195
i: [1,90,0,2,1,1,0,0] l3_polyred
a: [3,200,0,2,1,90]
s: 200
i: [1,89,0,2,1,1,0,0] l3_polyred
a: [3,205,0,2,1,89]
s: 205
i: [1,87,0,2,1,1,0,0] l3_polyred
a: [3,215,0,2,1,87]
s: 215
i: [1,85,0,2,1,1,0,0] l3_polyred
a: [3,225,0,2,1,85]
s: 225
i: [1,84,0,2,1,1,0,0] l3_polyred
a: [3,230,0,2,1,84]
s: 230
i: [1,99,0,3,1,1,0,0] l3_polyred
a: [3,155,0,3,1,99]
s: 155
i: [1,94,0,3,1,1,0,0] l3_polyred
a: [3,180,0,3,1,94]
s: 180
i: [1,92,0,3,1,1,0,0] l3_polyred
a: [3,190,0,3,1,92]
s: 190
i: [1,88,0,3,1,1,0,0] l3_polyred
a: [3,210,0,3,1,88]
s: 210
i: [1,86,0,3,1,1,0,0] l3_polyred
a: [3,220,0,3,1,86]
s: 220
i: [1,98,0,4,1,1,0,0] l3_polyred
a: [3,160,0,4,1,98]
s: 160
i: [1,97,0,4,1,1,0,0] l3_polyred
a: [3,165,0,4,1,97]
s: 165
i: [1,144,0,2,1,0,0,1] rc1_ordinal1
a: [3,235,0,2,1,144]
s: 235
i: [1,146,0,2,1,0,0,1] rc1_funct_1
a: [3,240,0,2,1,146]
s: 240
i: [1,145,0,2,1,0,0,1] rc1_funct_1
a: [3,245,0,2,1,145]
s: 245
i: [1,149,0,2,1,0,0,1] rc1_valued_0
a: [3,250,0,2,1,149]
s: 250
i: [1,148,0,2,1,0,0,1] rc1_valued_0
a: [3,255,0,2,1,148]
s: 255
i: [1,147,0,2,1,0,0,1] rc1_valued_0
a: [3,260,0,2,1,147]
s: 260
i: [1,120,0,3,1,0,0,1] rc6_valued_0
s: 120
i: [1,122,0,3,1,0,0,1] rc6_valued_0
s: 122
i: [1,124,0,3,1,0,0,1] rc6_valued_0
s: 124
i: [1,125,0,3,1,0,0,1] rc14_pre_poly
s: 125
i: [1,127,0,3,1,0,0,1] rc14_pre_poly
s: 127
i: [1,129,0,3,1,0,0,1] rc14_pre_poly
s: 129
i: [1,130,0,3,1,0,0,1] rc14_pre_poly
s: 130
i: [1,100,0,4,2,0,0,255] cc1_ordinal1
s: 100
i: [1,101,0,4,2,0,0,255] cc1_ordinal1
s: 101
d: [2,262,1,2,1,1,40,100,230]
a: [3,267,1,2,1,262]
s: 267
i: [1,121,0,4,1,0,0,1] rc6_valued_0
s: 121
d: [2,274,1,2,1,1,40,101,230]
a: [3,279,1,2,1,274]
s: 279
i: [1,123,0,4,1,0,0,1] rc6_valued_0
s: 123
i: [1,126,0,4,1,0,0,1] rc14_pre_poly
s: 126
d: [2,263,1,2,1,1,40,100,235]
a: [3,272,1,2,1,263]
s: 272
i: [1,128,0,4,1,0,0,1] rc14_pre_poly
s: 128
d: [2,275,1,2,1,1,40,101,235]
a: [3,284,1,2,1,275]
s: 284
i: [1,131,0,4,1,0,0,255] rc3_relat_1
s: 131
i: [1,134,0,4,1,0,0,255] rc6_funct_1
s: 134
i: [1,137,0,4,1,0,0,255] rc6_funct_1
s: 137
i: [1,132,0,5,1,0,0,255] rc3_relat_1
s: 132
i: [1,133,0,5,1,0,0,255] rc3_relat_1
s: 133
i: [1,135,0,5,1,0,0,255] rc6_funct_1
s: 135
i: [1,136,0,5,1,0,0,255] rc6_funct_1
s: 136
i: [1,105,0,6,3,0,0,255] cc2_ordinal1
s: 105
i: [1,109,0,6,3,0,0,255] cc8_valued_0
s: 109
d: [2,288,1,4,2,1,40,109,205]
d: [2,293,1,2,1,2,21,288,225]
a: [3,297,1,2,1,293]
s: 297
i: [1,111,0,7,3,0,0,255] cc7_valued_0
s: 111
i: [1,113,0,7,3,0,0,255] cc1_valued_0
s: 113
d: [2,289,1,4,2,1,40,109,175]
d: [2,299,1,2,1,2,21,289,195]
a: [3,303,1,2,1,299]
s: 303
i: [1,115,0,7,3,0,0,255] cc3_valued_0
s: 115
d: [2,292,1,4,2,1,40,109,250]
d: [2,307,1,2,1,2,21,292,260]
a: [3,311,1,2,1,307]
s: 311
i: [1,117,0,7,3,0,0,255] cc4_valued_0
s: 117
i: [1,119,0,8,3,0,0,255] cc2_valued_0
s: 119
d: [2,290,1,6,2,0,40,109,124]
d: [2,305,1,3,1,0,21,290,120]
s: 305
i: [1,150,0,24,9,0,0,1] redefinition_r6_pboole
d: [2,151,0,14,5,0,18,150]
s: 151
d: [2,291,1,6,2,0,40,109,129]
d: [2,306,1,3,1,0,21,291,125]
s: 306
i: [1,140,0,24,9,0,0,2] reflexivity_r6_pboole
s: 140
i: [1,142,0,27,10,0,0,1] redefinition_r6_pboole
s: 142
d: [2,314,1,8,2,0,40,115,133]
d: [2,317,1,4,1,0,21,314,131]
s: 317
i: [1,141,0,28,10,0,0,2] symmetry_r6_pboole
s: 141
d: [2,315,1,8,2,0,40,115,136]
d: [2,318,1,4,1,0,21,315,134]
s: 318
i: [1,138,0,34,14,0,0,1] reflexivity_r2_pre_poly
s: 138
i: [1,103,0,38,15,0,0,255] d10_pre_poly
s: 103
d: [2,320,1,8,2,0,40,117,133]
d: [2,323,1,4,1,0,21,320,131]
s: 323
i: [1,104,0,38,15,0,0,255] d10_pre_poly
s: 104
d: [2,321,1,8,2,0,40,117,136]
d: [2,324,1,4,1,0,21,321,134]
s: 324
i: [1,139,0,38,15,0,0,255] antisymmetry_r1_pre_poly
s: 139
i: [1,102,0,42,16,0,0,255] d10_pre_poly
s: 102
d: [2,337,1,38,15,1,40,102,160]
d: [2,341,1,36,14,2,21,337,200]
d: [2,342,1,34,13,3,21,341,205]
d: [2,343,1,31,12,4,21,342,210]
d: [2,344,1,29,11,5,21,343,215]
d: [2,345,1,26,10,6,21,344,220]
d: [2,346,1,24,9,7,21,345,225]
d: [2,347,1,22,8,8,21,346,170]
d: [2,348,1,20,7,9,21,347,175]
d: [2,349,1,17,6,10,21,348,180]
d: [2,350,1,15,5,11,21,349,185]
d: [2,351,1,12,4,12,21,350,190]
d: [2,352,1,10,3,13,21,351,195]
d: [2,353,1,8,2,14,21,352,230]
a: [3,361,1,4,1,353]
s: 361
d: [2,338,1,38,15,1,40,102,165]
d: [2,363,1,36,14,2,21,338,170]
d: [2,364,1,34,13,3,21,363,175]
d: [2,365,1,31,12,4,21,364,180]
d: [2,366,1,29,11,5,21,365,185]
d: [2,367,1,26,10,6,21,366,190]
d: [2,368,1,24,9,7,21,367,195]
d: [2,369,1,22,8,8,21,368,200]
d: [2,370,1,20,7,9,21,369,205]
d: [2,371,1,17,6,10,21,370,210]
d: [2,372,1,15,5,11,21,371,215]
d: [2,373,1,12,4,12,21,372,220]
d: [2,374,1,10,3,13,21,373,225]
d: [2,375,1,8,2,14,21,374,230]
a: [3,379,1,4,1,375]
s: 379
d: [2,401,2,0,0,10,52,225,195,215,185,220,210,190,180,155,379,142]
e: 401
d: [2,404,2,23,9,1,40,379,142]
d: [2,414,2,20,8,2,21,404,155]
d: [2,415,2,17,7,3,21,414,180]
d: [2,416,2,15,6,4,21,415,185]
d: [2,417,2,12,5,5,21,416,190]
d: [2,418,2,10,4,6,21,417,195]
d: [2,419,2,7,3,7,21,418,210]
d: [2,420,2,5,2,8,21,419,215]
d: [2,421,2,2,1,9,21,420,220]
d: [2,422,2,0,0,10,21,421,225]
e: 422
r: 361
r: 379
a: [3,357,1,4,1,353]
s: 357
d: [2,424,2,0,0,10,52,195,225,185,215,190,180,220,210,155,357,142]
e: 424
d: [2,427,2,23,9,1,40,357,142]
d: [2,437,2,20,8,2,21,427,155]
d: [2,438,2,17,7,3,21,437,210]
d: [2,439,2,15,6,4,21,438,215]
d: [2,440,2,12,5,5,21,439,220]
d: [2,441,2,10,4,6,21,440,225]
d: [2,442,2,7,3,7,21,441,180]
d: [2,443,2,5,2,8,21,442,185]
d: [2,444,2,2,1,9,21,443,190]
d: [2,445,2,0,0,10,21,444,195]
e: 445
d: [2,385,2,34,14,1,40,361,139]
d: [2,387,2,32,13,2,21,385,170]
d: [2,388,2,30,12,3,21,387,175]
d: [2,389,2,27,11,4,21,388,180]
d: [2,390,2,25,10,5,21,389,185]
d: [2,391,2,22,9,6,21,390,190]
d: [2,392,2,20,8,7,21,391,195]
d: [2,393,2,18,7,8,21,392,200]
d: [2,394,2,16,6,9,21,393,205]
d: [2,395,2,13,5,10,21,394,210]
d: [2,396,2,11,4,11,21,395,215]
d: [2,397,2,8,3,12,21,396,220]
d: [2,398,2,6,2,13,21,397,225]
d: [2,399,2,4,1,14,21,398,230]
f: 445,422,399,375,353,84,85,86,87,88,89,90,91,92,93,94,95,96,97,98,99
% Refutation found. Thanks to Tanya!
% SZS status Theorem for polyred__l3_polyred
% SZS output start Proof for polyred__l3_polyred
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : ((v2_pre_poly(X1) & v4_valued_0(X1) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pre_poly(X0,X2,X1) & r2_pre_poly(X0,X1,X2)) => X1 = X2)))) [input l3_polyred]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : ((v2_pre_poly(X1) & v4_valued_0(X1) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pre_poly(X0,X2,X1) & r2_pre_poly(X0,X1,X2)) => X1 = X2)))) [negated conjecture 1]
3. ! [X0,X1,X2] : ((v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v2_pre_poly(X1) & v4_valued_0(X1) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1) & v3_ordinal1(X0)) => (r1_pre_poly(X0,X1,X2) => ~r1_pre_poly(X0,X2,X1))) [input antisymmetry_r1_pre_poly]
14. ! [X0] : (v3_ordinal1(X0) => ! [X1] : ((v2_pre_poly(X1) & v4_valued_0(X1) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r2_pre_poly(X0,X1,X2) <=> (r6_pboole(X0,X1,X2) | r1_pre_poly(X0,X1,X2)))))) [input d10_pre_poly]
24. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (r6_pboole(X0,X1,X2) <=> X1 = X2)) [input redefinition_r6_pboole]
30. ? [X0] : (? [X1] : (? [X2] : ((X1 != X2 & (r2_pre_poly(X0,X2,X1) & r2_pre_poly(X0,X1,X2))) & (v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v2_pre_poly(X1) & v4_valued_0(X1) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) & v3_ordinal1(X0)) [ennf transformation 2]
31. ? [X0] : (? [X1] : (? [X2] : (X1 != X2 & r2_pre_poly(X0,X2,X1) & r2_pre_poly(X0,X1,X2) & v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v2_pre_poly(X1) & v4_valued_0(X1) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) & v3_ordinal1(X0)) [flattening 30]
33. ! [X0] : (! [X1] : (! [X2] : ((r2_pre_poly(X0,X1,X2) <=> (r6_pboole(X0,X1,X2) | r1_pre_poly(X0,X1,X2))) | (~v2_pre_poly(X2) | ~v4_valued_0(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v2_pre_poly(X1) | ~v4_valued_0(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) | ~v3_ordinal1(X0)) [ennf transformation 14]
34. ! [X0] : (! [X1] : (! [X2] : ((r2_pre_poly(X0,X1,X2) <=> (r6_pboole(X0,X1,X2) | r1_pre_poly(X0,X1,X2))) | ~v2_pre_poly(X2) | ~v4_valued_0(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v2_pre_poly(X1) | ~v4_valued_0(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) | ~v3_ordinal1(X0)) [flattening 33]
55. ! [X0,X1,X2] : ((~r1_pre_poly(X0,X2,X1) | ~r1_pre_poly(X0,X1,X2)) | (~v2_pre_poly(X2) | ~v4_valued_0(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v2_pre_poly(X1) | ~v4_valued_0(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v3_ordinal1(X0))) [ennf transformation 3]
56. ! [X0,X1,X2] : (~r1_pre_poly(X0,X2,X1) | ~r1_pre_poly(X0,X1,X2) | ~v2_pre_poly(X2) | ~v4_valued_0(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v2_pre_poly(X1) | ~v4_valued_0(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v3_ordinal1(X0)) [flattening 55]
61. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 24]
62. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 61]
63. ? [X0] : (? [X1] : (? [X2] : (X1 != X2 & r2_pre_poly(X0,X2,X1) & r2_pre_poly(X0,X1,X2) & v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v2_pre_poly(X1) & v4_valued_0(X1) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) & v3_ordinal1(X0)) => (? [X1] : (? [X2] : (X1 != X2 & r2_pre_poly(sK0,X2,X1) & r2_pre_poly(sK0,X1,X2) & v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v2_pre_poly(X1) & v4_valued_0(X1) & v1_partfun1(X1,sK0) & v1_funct_1(X1) & v4_relat_1(X1,sK0) & v1_relat_1(X1)) & v3_ordinal1(sK0)) [choice axiom]
64. ? [X1] : (? [X2] : (X1 != X2 & r2_pre_poly(sK0,X2,X1) & r2_pre_poly(sK0,X1,X2) & v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v2_pre_poly(X1) & v4_valued_0(X1) & v1_partfun1(X1,sK0) & v1_funct_1(X1) & v4_relat_1(X1,sK0) & v1_relat_1(X1)) => (? [X2] : (sK1 != X2 & r2_pre_poly(sK0,X2,sK1) & r2_pre_poly(sK0,sK1,X2) & v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v2_pre_poly(sK1) & v4_valued_0(sK1) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
65. ? [X2] : (sK1 != X2 & r2_pre_poly(sK0,X2,sK1) & r2_pre_poly(sK0,sK1,X2) & v2_pre_poly(X2) & v4_valued_0(X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (sK1 != sK2 & r2_pre_poly(sK0,sK2,sK1) & r2_pre_poly(sK0,sK1,sK2) & v2_pre_poly(sK2) & v4_valued_0(sK2) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
66. ((sK1 != sK2 & r2_pre_poly(sK0,sK2,sK1) & r2_pre_poly(sK0,sK1,sK2) & v2_pre_poly(sK2) & v4_valued_0(sK2) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v2_pre_poly(sK1) & v4_valued_0(sK1) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) & v3_ordinal1(sK0) [skolemisation 31,65,64,63]
67. ! [X0] : (! [X1] : (! [X2] : (((r2_pre_poly(X0,X1,X2) | (~r6_pboole(X0,X1,X2) & ~r1_pre_poly(X0,X1,X2))) & ((r6_pboole(X0,X1,X2) | r1_pre_poly(X0,X1,X2)) | ~r2_pre_poly(X0,X1,X2))) | ~v2_pre_poly(X2) | ~v4_valued_0(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v2_pre_poly(X1) | ~v4_valued_0(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) | ~v3_ordinal1(X0)) [nnf transformation 34]
68. ! [X0] : (! [X1] : (! [X2] : (((r2_pre_poly(X0,X1,X2) | (~r6_pboole(X0,X1,X2) & ~r1_pre_poly(X0,X1,X2))) & (r6_pboole(X0,X1,X2) | r1_pre_poly(X0,X1,X2) | ~r2_pre_poly(X0,X1,X2))) | ~v2_pre_poly(X2) | ~v4_valued_0(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v2_pre_poly(X1) | ~v4_valued_0(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) | ~v3_ordinal1(X0)) [flattening 67]
77. ! [X0,X1,X2] : (((r6_pboole(X0,X1,X2) | X1 != X2) & (X1 = X2 | ~r6_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 62]
84. v3_ordinal1(sK0) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
85. v1_relat_1(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
86. v4_relat_1(sK1,sK0) [cnf transformation 66] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
87. v1_funct_1(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
88. v1_partfun1(sK1,sK0) [cnf transformation 66] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
89. v4_valued_0(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
90. v2_pre_poly(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
91. v1_relat_1(sK2) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
92. v4_relat_1(sK2,sK0) [cnf transformation 66] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
93. v1_funct_1(sK2) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
94. v1_partfun1(sK2,sK0) [cnf transformation 66] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
95. v4_valued_0(sK2) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
96. v2_pre_poly(sK2) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
97. r2_pre_poly(sK0,sK1,sK2) [cnf transformation 66] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
98. r2_pre_poly(sK0,sK2,sK1) [cnf transformation 66] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
99. sK1 != sK2 [cnf transformation 66] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
102. ~r2_pre_poly(X0,X1,X2) | r1_pre_poly(X0,X1,X2) | r6_pboole(X0,X1,X2) | ~v2_pre_poly(X2) | ~v4_valued_0(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v2_pre_poly(X1) | ~v4_valued_0(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v3_ordinal1(X0) [cnf transformation 68] {a:0,w:42,wCS:16800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
139. ~r1_pre_poly(X0,X2,X1) | ~r1_pre_poly(X0,X1,X2) | ~v2_pre_poly(X2) | ~v4_valued_0(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v2_pre_poly(X1) | ~v4_valued_0(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v3_ordinal1(X0) [cnf transformation 56] {a:0,w:38,wCS:15200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
142. ~r6_pboole(X0,X1,X2) | X1 = X2 | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 77] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
153. 1 <=> sK1 = sK2 [avatar definition]
155. sK1 != sK2 <- (~1) [avatar component clause 153] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
156. ~1 [avatar split clause 99,153]
158. 2 <=> r2_pre_poly(sK0,sK2,sK1) [avatar definition]
160. r2_pre_poly(sK0,sK2,sK1) <- (2) [avatar component clause 158] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
161. 2 [avatar split clause 98,158]
163. 3 <=> r2_pre_poly(sK0,sK1,sK2) [avatar definition]
165. r2_pre_poly(sK0,sK1,sK2) <- (3) [avatar component clause 163] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
166. 3 [avatar split clause 97,163]
168. 4 <=> v2_pre_poly(sK2) [avatar definition]
170. v2_pre_poly(sK2) <- (4) [avatar component clause 168] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
171. 4 [avatar split clause 96,168]
173. 5 <=> v4_valued_0(sK2) [avatar definition]
175. v4_valued_0(sK2) <- (5) [avatar component clause 173] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
176. 5 [avatar split clause 95,173]
178. 6 <=> v1_partfun1(sK2,sK0) [avatar definition]
180. v1_partfun1(sK2,sK0) <- (6) [avatar component clause 178] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
181. 6 [avatar split clause 94,178]
183. 7 <=> v1_funct_1(sK2) [avatar definition]
185. v1_funct_1(sK2) <- (7) [avatar component clause 183] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
186. 7 [avatar split clause 93,183]
188. 8 <=> v4_relat_1(sK2,sK0) [avatar definition]
190. v4_relat_1(sK2,sK0) <- (8) [avatar component clause 188] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
191. 8 [avatar split clause 92,188]
193. 9 <=> v1_relat_1(sK2) [avatar definition]
195. v1_relat_1(sK2) <- (9) [avatar component clause 193] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
196. 9 [avatar split clause 91,193]
198. 10 <=> v2_pre_poly(sK1) [avatar definition]
200. v2_pre_poly(sK1) <- (10) [avatar component clause 198] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
201. 10 [avatar split clause 90,198]
203. 11 <=> v4_valued_0(sK1) [avatar definition]
205. v4_valued_0(sK1) <- (11) [avatar component clause 203] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
206. 11 [avatar split clause 89,203]
208. 12 <=> v1_partfun1(sK1,sK0) [avatar definition]
210. v1_partfun1(sK1,sK0) <- (12) [avatar component clause 208] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
211. 12 [avatar split clause 88,208]
213. 13 <=> v1_funct_1(sK1) [avatar definition]
215. v1_funct_1(sK1) <- (13) [avatar component clause 213] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
216. 13 [avatar split clause 87,213]
218. 14 <=> v4_relat_1(sK1,sK0) [avatar definition]
220. v4_relat_1(sK1,sK0) <- (14) [avatar component clause 218] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
221. 14 [avatar split clause 86,218]
223. 15 <=> v1_relat_1(sK1) [avatar definition]
225. v1_relat_1(sK1) <- (15) [avatar component clause 223] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
226. 15 [avatar split clause 85,223]
228. 16 <=> v3_ordinal1(sK0) [avatar definition]
230. v3_ordinal1(sK0) <- (16) [avatar component clause 228] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
231. 16 [avatar split clause 84,228]
337. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v2_pre_poly(sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2) [resolution 102,160] {a:1,w:38,wCS:3800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
338. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v2_pre_poly(sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3) [resolution 102,165] {a:1,w:38,wCS:3800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
341. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v2_pre_poly(sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 10) [subsumption resolution 337,200] {a:1,w:36,wCS:3600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
342. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v2_pre_poly(sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 10, 11) [subsumption resolution 341,205] {a:1,w:34,wCS:3400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
343. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v2_pre_poly(sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 10, 11, 12) [subsumption resolution 342,210] {a:1,w:31,wCS:3100,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
344. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v2_pre_poly(sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 10, 11, 12, 13) [subsumption resolution 343,215] {a:1,w:29,wCS:2900,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
345. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v1_relat_1(sK1) | ~v2_pre_poly(sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 10, 11, 12, 13, 14) [subsumption resolution 344,220] {a:1,w:26,wCS:2600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
346. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v2_pre_poly(sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 10, 11, 12, 13, 14, 15) [subsumption resolution 345,225] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
347. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 4, 10, 11, 12, 13, 14, 15) [subsumption resolution 346,170] {a:1,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
348. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 4, 5, 10, 11, 12, 13, 14, 15) [subsumption resolution 347,175] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
349. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 4, 5, 6, 10, 11, 12, 13, 14, 15) [subsumption resolution 348,180] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
350. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 4, 5, 6, 7, 10, 11, 12, 13, 14, 15) [subsumption resolution 349,185] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
351. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v1_relat_1(sK2) | ~v3_ordinal1(sK0) <- (2, 4, 5, 6, 7, 8, 10, 11, 12, 13, 14, 15) [subsumption resolution 350,190] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
352. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) | ~v3_ordinal1(sK0) <- (2, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15) [subsumption resolution 351,195] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
353. r1_pre_poly(sK0,sK2,sK1) | r6_pboole(sK0,sK2,sK1) <- (2, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16) [subsumption resolution 352,230] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
355. 30 <=> r6_pboole(sK0,sK2,sK1) [avatar definition]
357. r6_pboole(sK0,sK2,sK1) <- (30) [avatar component clause 355] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
359. 31 <=> r1_pre_poly(sK0,sK2,sK1) [avatar definition]
361. r1_pre_poly(sK0,sK2,sK1) <- (31) [avatar component clause 359] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
362. 30 | 31 | ~2 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | ~15 | ~16 [avatar split clause 353,228,223,218,213,208,203,198,193,188,183,178,173,168,158,359,355]
363. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4) [subsumption resolution 338,170] {a:1,w:36,wCS:3600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
364. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5) [subsumption resolution 363,175] {a:1,w:34,wCS:3400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
365. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6) [subsumption resolution 364,180] {a:1,w:31,wCS:3100,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
366. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6, 7) [subsumption resolution 365,185] {a:1,w:29,wCS:2900,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
367. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6, 7, 8) [subsumption resolution 366,190] {a:1,w:26,wCS:2600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
368. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6, 7, 8, 9) [subsumption resolution 367,195] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
369. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6, 7, 8, 9, 10) [subsumption resolution 368,200] {a:1,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
370. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6, 7, 8, 9, 10, 11) [subsumption resolution 369,205] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
371. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6, 7, 8, 9, 10, 11, 12) [subsumption resolution 370,210] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
372. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13) [subsumption resolution 371,215] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
373. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14) [subsumption resolution 372,220] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
374. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) | ~v3_ordinal1(sK0) <- (3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15) [subsumption resolution 373,225] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
375. r1_pre_poly(sK0,sK1,sK2) | r6_pboole(sK0,sK1,sK2) <- (3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16) [subsumption resolution 374,230] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
377. 32 <=> r6_pboole(sK0,sK1,sK2) [avatar definition]
379. r6_pboole(sK0,sK1,sK2) <- (32) [avatar component clause 377] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
381. 33 <=> r1_pre_poly(sK0,sK1,sK2) [avatar definition]
384. 32 | 33 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | ~15 | ~16 [avatar split clause 375,228,223,218,213,208,203,198,193,188,183,178,173,168,163,381,377]
385. ~r1_pre_poly(sK0,sK1,sK2) | ~v2_pre_poly(sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (31) [resolution 361,139] {a:2,w:34,wCS:3400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
387. ~r1_pre_poly(sK0,sK1,sK2) | ~v4_valued_0(sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 31) [subsumption resolution 385,170] {a:2,w:32,wCS:3200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
388. ~r1_pre_poly(sK0,sK1,sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 31) [subsumption resolution 387,175] {a:2,w:30,wCS:3000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
389. ~r1_pre_poly(sK0,sK1,sK2) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 6, 31) [subsumption resolution 388,180] {a:2,w:27,wCS:2700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
390. ~r1_pre_poly(sK0,sK1,sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 6, 7, 31) [subsumption resolution 389,185] {a:2,w:25,wCS:2500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
391. ~r1_pre_poly(sK0,sK1,sK2) | ~v1_relat_1(sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 6, 7, 8, 31) [subsumption resolution 390,190] {a:2,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
392. ~r1_pre_poly(sK0,sK1,sK2) | ~v2_pre_poly(sK1) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 6, 7, 8, 9, 31) [subsumption resolution 391,195] {a:2,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
393. ~r1_pre_poly(sK0,sK1,sK2) | ~v4_valued_0(sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 6, 7, 8, 9, 10, 31) [subsumption resolution 392,200] {a:2,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
394. ~r1_pre_poly(sK0,sK1,sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 6, 7, 8, 9, 10, 11, 31) [subsumption resolution 393,205] {a:2,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
395. ~r1_pre_poly(sK0,sK1,sK2) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 6, 7, 8, 9, 10, 11, 12, 31) [subsumption resolution 394,210] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
396. ~r1_pre_poly(sK0,sK1,sK2) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 31) [subsumption resolution 395,215] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
397. ~r1_pre_poly(sK0,sK1,sK2) | ~v1_relat_1(sK1) | ~v3_ordinal1(sK0) <- (4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 31) [subsumption resolution 396,220] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
398. ~r1_pre_poly(sK0,sK1,sK2) | ~v3_ordinal1(sK0) <- (4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 31) [subsumption resolution 397,225] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
399. ~r1_pre_poly(sK0,sK1,sK2) <- (4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 31) [subsumption resolution 398,230] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
400. ~33 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | ~15 | ~16 | ~31 [avatar split clause 399,359,228,223,218,213,208,203,198,193,188,183,178,173,168,381]
404. sK1 = sK2 | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (32) [resolution 379,142] {a:2,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
414. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 32) [subsumption resolution 404,155] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
415. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 6, 32) [subsumption resolution 414,180] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
416. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 6, 7, 32) [subsumption resolution 415,185] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
417. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 6, 7, 8, 32) [subsumption resolution 416,190] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
418. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 6, 7, 8, 9, 32) [subsumption resolution 417,195] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
419. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 6, 7, 8, 9, 12, 32) [subsumption resolution 418,210] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
420. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 6, 7, 8, 9, 12, 13, 32) [subsumption resolution 419,215] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
421. ~v1_relat_1(sK1) <- (~1, 6, 7, 8, 9, 12, 13, 14, 32) [subsumption resolution 420,220] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
422. $false <- (~1, 6, 7, 8, 9, 12, 13, 14, 15, 32) [subsumption resolution 421,225] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
423. 1 | ~6 | ~7 | ~8 | ~9 | ~12 | ~13 | ~14 | ~15 | ~32 [avatar contradiction clause 422]
427. sK1 = sK2 | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (30) [resolution 357,142] {a:2,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
437. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 30) [subsumption resolution 427,155] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
438. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 12, 30) [subsumption resolution 437,210] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
439. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 12, 13, 30) [subsumption resolution 438,215] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
440. ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 12, 13, 14, 30) [subsumption resolution 439,220] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
441. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 12, 13, 14, 15, 30) [subsumption resolution 440,225] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
442. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 6, 12, 13, 14, 15, 30) [subsumption resolution 441,180] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
443. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 6, 7, 12, 13, 14, 15, 30) [subsumption resolution 442,185] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
444. ~v1_relat_1(sK2) <- (~1, 6, 7, 8, 12, 13, 14, 15, 30) [subsumption resolution 443,190] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
445. $false <- (~1, 6, 7, 8, 9, 12, 13, 14, 15, 30) [subsumption resolution 444,195] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
446. 1 | ~6 | ~7 | ~8 | ~9 | ~12 | ~13 | ~14 | ~15 | ~30 [avatar contradiction clause 445]
447. $false [avatar sat refutation 156,161,166,171,176,181,186,191,196,201,206,211,216,221,226,231,362,384,400,423,446] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:303,thDist:-303}
% SZS output end Proof for polyred__l3_polyred
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 66
% Generated clauses: 230
% Active clauses: 73
% Passive clauses: 81
% Final active clauses: 71
% Final passive clauses: 6
% Main loop iterations started: 74
% 
% >>> Simplifying Inferences
% Duplicate literals: 46
% Fw subsumption resolutions: 84
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Equational tautologies: 1
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 34
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 38
% SAT solver unit clauses: 22
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward superposition index maintenance: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

