i: [1,91,0,3,1,1,0,0] t6_funct_1
a: [3,134,0,3,1,91]
s: 134
i: [1,92,0,2,1,0,0,255] spc1_boole
a: [3,139,0,2,1,92]
s: 139
i: [1,93,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,144,0,2,1,93]
s: 144
i: [1,94,0,2,1,0,0,1] fc1_xboole_0
a: [3,149,0,2,1,94]
s: 149
i: [1,121,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,159,0,2,1,121]
s: 159
i: [1,122,0,2,1,0,0,2] rc2_xboole_0
a: [3,164,0,2,1,122]
s: 164
i: [1,123,0,2,1,0,0,2] rc1_xboole_0
a: [3,169,0,2,1,123]
s: 169
i: [1,125,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,174,0,2,1,125]
s: 174
i: [1,124,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,179,0,2,1,124]
s: 179
i: [1,127,0,2,1,0,0,1] rc1_relat_1
a: [3,184,0,2,1,127]
s: 184
i: [1,126,0,2,1,0,0,1] rc1_relat_1
a: [3,189,0,2,1,126]
s: 189
i: [1,128,0,2,1,0,0,255] rc2_relat_1
a: [3,194,0,2,1,128]
s: 194
i: [1,130,0,2,1,0,0,1] rc1_funct_1
a: [3,199,0,2,1,130]
s: 199
i: [1,129,0,2,1,0,0,1] rc1_funct_1
a: [3,204,0,2,1,129]
s: 204
i: [1,112,0,3,1,0,0,255] s3_funct_1__e2_15__funct_1
s: 112
i: [1,113,0,3,1,0,0,255] s3_funct_1__e2_15__funct_1
s: 113
i: [1,116,0,3,1,0,0,255] s3_funct_1__e7_15__funct_1
s: 116
i: [1,117,0,3,1,0,0,255] s3_funct_1__e7_15__funct_1
s: 117
i: [1,95,0,3,1,0,0,255] d2_xboole_0
a: [3,154,0,3,1,95]
s: 154
r: 144
i: [1,96,0,4,1,0,0,255] dt_o_1_2_funct_1
s: 96
i: [1,97,0,4,2,0,0,255] cc2_zfmisc_1
s: 97
d: [2,207,1,2,1,1,40,97,169]
a: [3,212,1,2,1,207]
s: 212
i: [1,98,0,4,2,0,0,255] cc1_zfmisc_1
s: 98
d: [2,208,1,2,1,1,40,97,149]
a: [3,217,1,2,1,208]
s: 217
i: [1,99,0,4,2,0,0,1] cc1_funct_1
s: 99
i: [1,100,0,4,2,0,0,1] cc1_relat_1
s: 100
d: [2,219,1,2,1,1,40,99,169]
a: [3,224,1,2,1,219]
s: 224
i: [1,111,0,4,1,0,0,255] existence_m1_subset_1
s: 111
d: [2,220,1,2,1,1,40,99,149]
a: [3,229,1,2,1,220]
s: 229
i: [1,90,0,19,7,1,0,0] t6_funct_1
s: 90
i: [1,101,0,5,2,0,0,1] t6_boole
s: 101
d: [2,231,1,2,1,1,40,100,169]
a: [3,236,1,2,1,231]
s: 236
i: [1,102,0,5,2,0,0,1] fc10_relat_1
s: 102
d: [2,232,1,2,1,1,40,100,149]
a: [3,241,1,2,1,232]
s: 241
i: [1,110,0,5,2,0,0,255] t7_boole
s: 110
i: [1,114,0,5,1,0,0,255] s3_funct_1__e2_15__funct_1
s: 114
d: [2,252,1,21,7,0,44,90,114]
d: [2,253,1,18,6,0,21,252,113]
d: [2,254,1,15,5,0,21,253,112]
s: 254
i: [1,118,0,5,1,0,0,255] s3_funct_1__e7_15__funct_1
s: 118
d: [2,255,2,17,5,0,44,254,114]
d: [2,256,2,14,4,0,21,255,113]
d: [2,257,2,11,3,0,21,256,112]
s: 257
i: [1,107,0,6,2,0,0,255] t1_subset
s: 107
i: [1,120,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 120
d: [2,265,3,8,2,0,47,257]
s: 265
i: [1,103,0,7,3,0,0,1] fc8_relat_1
s: 103
d: [2,258,2,17,5,0,44,254,118]
d: [2,261,2,14,4,0,21,258,117]
d: [2,262,2,11,3,0,21,261,116]
s: 262
i: [1,104,0,7,3,0,0,255] fc24_relat_1
s: 104
i: [1,109,0,7,3,0,0,1] t8_boole
s: 109
d: [2,272,3,8,2,0,47,262]
s: 272
i: [1,108,0,8,3,0,0,255] t2_subset
s: 108
d: [2,277,4,5,1,0,47,272]
a: [3,281,4,5,1,277]
s: 281
i: [1,115,0,9,2,0,0,255] s3_funct_1__e2_15__funct_1
s: 115
i: [1,119,0,9,2,0,0,255] s3_funct_1__e7_15__funct_1
s: 119
d: [2,288,5,9,2,1,44,119,281]
s: 288
d: [2,260,1,21,7,0,44,90,118]
d: [2,263,1,18,6,0,21,260,117]
d: [2,264,1,15,5,0,21,263,116]
s: 264
d: [2,289,6,9,3,1,44,288,115]
d: [2,292,6,6,2,1,18,289]
a: [3,295,6,3,1,292]
s: 295
r: 288
d: [2,247,1,5,2,0,40,102,100]
s: 247
d: [2,248,1,5,2,0,40,102,99]
s: 248
a: [3,298,6,3,1,292]
s: 298
d: [2,249,1,5,2,0,40,102,97]
s: 249
d: [2,304,2,17,5,0,44,264,118]
d: [2,305,2,14,4,0,21,304,117]
d: [2,306,2,11,3,0,21,305,116]
s: 306
d: [2,251,1,5,2,0,44,102,114]
s: 251
d: [2,259,1,5,2,0,44,102,118]
s: 259
d: [2,268,1,8,3,0,44,103,114]
d: [2,270,1,5,2,0,21,268,112]
s: 270
d: [2,269,1,8,3,0,44,103,118]
d: [2,271,1,5,2,0,21,269,116]
s: 271
d: [2,273,1,8,3,0,44,104,114]
d: [2,275,1,5,2,0,21,273,112]
s: 275
d: [2,274,1,8,3,0,44,104,118]
d: [2,276,1,5,2,0,21,274,116]
s: 276
d: [2,283,1,6,2,0,40,108,96]
s: 283
d: [2,329,7,2,1,1,40,283,295]
a: [3,333,7,2,1,329]
s: 333
d: [2,335,8,0,0,2,52,134,333,101]
e: 335
d: [2,336,8,0,0,3,52,149,134,333,109]
e: 336
d: [2,337,8,0,0,3,52,149,134,333,109]
e: 337
d: [2,290,6,9,3,1,44,115,288]
d: [2,291,6,6,2,1,18,290]
f: 335,329,291,277,95,93,92,91
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_1__t6_funct_1
% SZS output start Proof for funct_1__t6_funct_1
1. ! [X0] : (! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((k9_xtuple_0(X2) = X0 & k9_xtuple_0(X1) = X0) => X1 = X2))) => k1_xboole_0 = X0) [input t6_funct_1]
2. ~! [X0] : (! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((k9_xtuple_0(X2) = X0 & k9_xtuple_0(X1) = X0) => X1 = X2))) => k1_xboole_0 = X0) [negated conjecture 1]
10. k1_xboole_0 = o_0_0_xboole_0 [input d2_xboole_0]
15. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
16. ! [X0] : m1_subset_1(o_1_2_funct_1(X0),X0) [input dt_o_1_2_funct_1]
29. ! [X0] : ? [X1] : (! [X2] : (r2_hidden(X2,X0) => k1_xboole_0 = k1_funct_1(X1,X2)) & k9_xtuple_0(X1) = X0 & v1_funct_1(X1) & v1_relat_1(X1)) [input s3_funct_1__e2_15__funct_1]
30. ! [X0] : ? [X1] : (! [X2] : (r2_hidden(X2,X0) => k1_funct_1(X1,X2) = 1) & k9_xtuple_0(X1) = X0 & v1_funct_1(X1) & v1_relat_1(X1)) [input s3_funct_1__e7_15__funct_1]
31. ~v1_xboole_0(1) [input spc1_boole]
33. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
34. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
38. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 33]
44. ? [X0] : (k1_xboole_0 != X0 & ! [X1] : (! [X2] : ((X1 = X2 | (k9_xtuple_0(X2) != X0 | k9_xtuple_0(X1) != X0)) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1)))) [ennf transformation 2]
45. ? [X0] : (k1_xboole_0 != X0 & ! [X1] : (! [X2] : (X1 = X2 | k9_xtuple_0(X2) != X0 | k9_xtuple_0(X1) != X0 | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1))) [flattening 44]
50. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 34]
61. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 38]
62. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 61]
65. ! [X0] : ? [X1] : (! [X2] : (k1_xboole_0 = k1_funct_1(X1,X2) | ~r2_hidden(X2,X0)) & k9_xtuple_0(X1) = X0 & v1_funct_1(X1) & v1_relat_1(X1)) [ennf transformation 29]
66. ! [X0] : ? [X1] : (! [X2] : (k1_funct_1(X1,X2) = 1 | ~r2_hidden(X2,X0)) & k9_xtuple_0(X1) = X0 & v1_funct_1(X1) & v1_relat_1(X1)) [ennf transformation 30]
68. ? [X0] : (k1_xboole_0 != X0 & ! [X1] : (! [X2] : (X1 = X2 | k9_xtuple_0(X2) != X0 | k9_xtuple_0(X1) != X0 | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1))) => (k1_xboole_0 != sK0 & ! [X1] : (! [X2] : (X1 = X2 | k9_xtuple_0(X2) != sK0 | k9_xtuple_0(X1) != sK0 | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1))) [choice axiom]
69. k1_xboole_0 != sK0 & ! [X1] : (! [X2] : (X1 = X2 | k9_xtuple_0(X2) != sK0 | k9_xtuple_0(X1) != sK0 | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) [skolemisation 45,68]
72. ! [X0] : (? [X1] : (! [X2] : (k1_xboole_0 = k1_funct_1(X1,X2) | ~r2_hidden(X2,X0)) & k9_xtuple_0(X1) = X0 & v1_funct_1(X1) & v1_relat_1(X1)) => (! [X2] : (k1_xboole_0 = k1_funct_1(sK2(X0),X2) | ~r2_hidden(X2,X0)) & k9_xtuple_0(sK2(X0)) = X0 & v1_funct_1(sK2(X0)) & v1_relat_1(sK2(X0)))) [choice axiom]
73. ! [X0] : (! [X2] : (k1_xboole_0 = k1_funct_1(sK2(X0),X2) | ~r2_hidden(X2,X0)) & k9_xtuple_0(sK2(X0)) = X0 & v1_funct_1(sK2(X0)) & v1_relat_1(sK2(X0))) [skolemisation 65,72]
74. ! [X0] : (? [X1] : (! [X2] : (k1_funct_1(X1,X2) = 1 | ~r2_hidden(X2,X0)) & k9_xtuple_0(X1) = X0 & v1_funct_1(X1) & v1_relat_1(X1)) => (! [X2] : (1 = k1_funct_1(sK3(X0),X2) | ~r2_hidden(X2,X0)) & k9_xtuple_0(sK3(X0)) = X0 & v1_funct_1(sK3(X0)) & v1_relat_1(sK3(X0)))) [choice axiom]
75. ! [X0] : (! [X2] : (1 = k1_funct_1(sK3(X0),X2) | ~r2_hidden(X2,X0)) & k9_xtuple_0(sK3(X0)) = X0 & v1_funct_1(sK3(X0)) & v1_relat_1(sK3(X0))) [skolemisation 66,74]
90. k9_xtuple_0(X2) != sK0 | X1 = X2 | k9_xtuple_0(X1) != sK0 | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [cnf transformation 69] {a:0,w:19,wCS:1900,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
91. k1_xboole_0 != sK0 [cnf transformation 69] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
92. ~v1_xboole_0(1) [cnf transformation 31] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
93. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 15] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
95. k1_xboole_0 = o_0_0_xboole_0 [cnf transformation 10] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
96. m1_subset_1(o_1_2_funct_1(X0),X0) [cnf transformation 16] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
101. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 50] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
108. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 62] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
112. v1_relat_1(sK2(X0)) [cnf transformation 73] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
113. v1_funct_1(sK2(X0)) [cnf transformation 73] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
114. k9_xtuple_0(sK2(X0)) = X0 [cnf transformation 73] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
115. k1_xboole_0 = k1_funct_1(sK2(X0),X2) | ~r2_hidden(X2,X0) [cnf transformation 73] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. v1_relat_1(sK3(X0)) [cnf transformation 75] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
117. v1_funct_1(sK3(X0)) [cnf transformation 75] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
118. k9_xtuple_0(sK3(X0)) = X0 [cnf transformation 75] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
119. 1 = k1_funct_1(sK3(X0),X2) | ~r2_hidden(X2,X0) [cnf transformation 75] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
132. 1 <=> k1_xboole_0 = sK0 [avatar definition]
134. k1_xboole_0 != sK0 <- (~1) [avatar component clause 132] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
135. ~1 [avatar split clause 91,132]
137. 2 <=> v1_xboole_0(1) [avatar definition]
140. ~2 [avatar split clause 92,137]
142. 3 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
145. 3 [avatar split clause 93,142]
152. 5 <=> k1_xboole_0 = o_0_0_xboole_0 [avatar definition]
155. 5 [avatar split clause 95,152]
252. sK0 != X1 | sK2(X1) = X2 | k9_xtuple_0(X2) != sK0 | ~v1_funct_1(sK2(X1)) | ~v1_relat_1(sK2(X1)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [superposition 90,114] {a:1,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
253. sK0 != X1 | sK2(X1) = X2 | k9_xtuple_0(X2) != sK0 | ~v1_relat_1(sK2(X1)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [subsumption resolution 252,113] {a:1,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
254. k9_xtuple_0(X2) != sK0 | sK2(X1) = X2 | sK0 != X1 | ~v1_funct_1(X2) | ~v1_relat_1(X2) [subsumption resolution 253,112] {a:1,w:15,wCS:1500,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
258. sK0 != X0 | sK3(X0) = sK2(X1) | sK0 != X1 | ~v1_funct_1(sK3(X0)) | ~v1_relat_1(sK3(X0)) [superposition 254,118] {a:2,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
261. sK0 != X0 | sK3(X0) = sK2(X1) | sK0 != X1 | ~v1_relat_1(sK3(X0)) [subsumption resolution 258,117] {a:2,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
262. sK0 != X1 | sK3(X0) = sK2(X1) | sK0 != X0 [subsumption resolution 261,116] {a:2,w:11,wCS:1100,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
272. sK0 != X0 | sK3(X0) = sK2(sK0) [equality resolution 262] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
277. sK2(sK0) = sK3(sK0) [equality resolution 272] {a:4,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
279. 22 <=> sK2(sK0) = sK3(sK0) [avatar definition]
281. sK2(sK0) = sK3(sK0) <- (22) [avatar component clause 279] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
282. 22 [avatar split clause 277,279]
283. r2_hidden(o_1_2_funct_1(X0),X0) | v1_xboole_0(X0) [resolution 108,96] {a:1,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
288. 1 = k1_funct_1(sK2(sK0),X0) | ~r2_hidden(X0,sK0) <- (22) [superposition 119,281] {a:5,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
290. k1_xboole_0 = 1 | ~r2_hidden(X0,sK0) | ~r2_hidden(X0,sK0) <- (22) [superposition 115,288] {a:6,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
291. k1_xboole_0 = 1 | ~r2_hidden(X0,sK0) <- (22) [duplicate literal removal 290] {a:6,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
294. 23 <=> ! [X0] : ~r2_hidden(X0,sK0) [avatar definition]
295. ~r2_hidden(X0,sK0) <- (23) [avatar component clause 294] {a:6,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
297. 24 <=> k1_xboole_0 = 1 [avatar definition]
301. 23 | 24 | ~22 [avatar split clause 291,279,297,294]
302. k1_xboole_0 != 1 | k1_xboole_0 != o_0_0_xboole_0 | v1_xboole_0(1) | ~v1_xboole_0(o_0_0_xboole_0) [theory tautology sat conflict] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:1,thDist:-1}
329. v1_xboole_0(sK0) <- (23) [resolution 283,295] {a:7,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:12,thDist:-12}
331. 25 <=> v1_xboole_0(sK0) [avatar definition]
333. v1_xboole_0(sK0) <- (25) [avatar component clause 331] {a:7,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:12,thDist:-12}
334. 25 | ~23 [avatar split clause 329,294,331]
335. $false <- (~1, 25) [unit resulting resolution 134,333,101] {a:8,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
341. 1 | ~25 [avatar contradiction clause 335]
344. $false [avatar sat refutation 135,140,145,155,282,301,302,334,341] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:163,thDist:-163}
% SZS output end Proof for funct_1__t6_funct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 35
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 151
% Active clauses: 66
% Passive clauses: 78
% Final active clauses: 64
% Final passive clauses: 11
% Main loop iterations started: 67
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 14
% Fw demodulations: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Equational tautologies: 1
% Forward subsumptions: 18
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 27
% Unit resulting resolution: 3
% Forward superposition: 19
% Backward superposition: 12
% Self superposition: 1
% Equality resolution: 5
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 31
% SAT solver unit clauses: 16
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

