i: [1,113,0,10,4,1,0,0] t2_pzfmisc1
a: [3,172,0,2,1,113]
s: 172
a: [3,180,0,2,1,113]
s: 180
a: [3,177,0,3,1,113]
s: 177
i: [1,112,0,3,1,1,0,0] t2_pzfmisc1
a: [3,190,0,3,1,112]
s: 190
r: 177
r: 180
r: 172
d: [2,300,0,2,1,2,23,180,190]
a: [3,310,0,2,1,300]
s: 310
d: [2,301,0,2,1,2,23,172,190]
a: [3,315,0,2,1,301]
s: 315
d: [2,299,0,3,1,2,23,177,190]
a: [3,305,0,3,1,299]
s: 305
i: [1,114,0,2,1,0,0,1] fc1_xboole_0
a: [3,197,0,2,1,114]
s: 197
i: [1,152,0,2,1,0,0,2] rc2_xboole_0
a: [3,212,0,2,1,152]
s: 212
i: [1,153,0,2,1,0,0,2] rc1_xboole_0
a: [3,217,0,2,1,153]
s: 217
i: [1,157,0,2,1,0,0,255] rc9_finset_1
a: [3,222,0,2,1,157]
s: 222
i: [1,156,0,2,1,0,0,255] rc9_finset_1
a: [3,227,0,2,1,156]
s: 227
i: [1,155,0,2,1,0,0,255] rc9_finset_1
a: [3,232,0,2,1,155]
s: 232
i: [1,154,0,2,1,0,0,255] rc9_finset_1
a: [3,237,0,2,1,154]
s: 237
i: [1,161,0,2,1,0,0,1] rc3_finset_1
a: [3,242,0,2,1,161]
s: 242
i: [1,160,0,2,1,0,0,1] rc3_finset_1
a: [3,247,0,2,1,160]
s: 247
i: [1,159,0,2,1,0,0,1] rc3_finset_1
a: [3,252,0,2,1,159]
s: 252
i: [1,158,0,2,1,0,0,1] rc3_finset_1
a: [3,257,0,2,1,158]
s: 257
i: [1,163,0,2,1,0,0,2] rc1_finset_1
a: [3,262,0,2,1,163]
s: 262
i: [1,162,0,2,1,0,0,2] rc1_finset_1
a: [3,267,0,2,1,162]
s: 267
i: [1,166,0,2,1,0,0,255] rc8_finset_1
a: [3,272,0,2,1,166]
s: 272
i: [1,165,0,2,1,0,0,255] rc8_finset_1
a: [3,277,0,2,1,165]
s: 277
i: [1,164,0,2,1,0,0,255] rc8_finset_1
a: [3,282,0,2,1,164]
s: 282
i: [1,169,0,2,1,0,0,255] rc1_funcop_1
a: [3,287,0,2,1,169]
s: 287
i: [1,168,0,2,1,0,0,255] rc1_funcop_1
a: [3,292,0,2,1,168]
s: 292
i: [1,167,0,2,1,0,0,255] rc1_funcop_1
a: [3,297,0,2,1,167]
s: 297
i: [1,149,0,3,1,0,0,255] reflexivity_r1_tarski
s: 149
i: [1,123,0,4,2,0,0,255] cc5_finset_1
s: 123
i: [1,124,0,4,2,0,0,255] cc4_finset_1
s: 124
i: [1,125,0,4,2,0,0,255] cc6_finset_1
s: 125
i: [1,126,0,4,2,0,0,2] cc1_finset_1
s: 126
d: [2,317,1,2,1,1,40,125,197]
a: [3,322,1,2,1,317]
s: 322
i: [1,148,0,4,1,0,0,3] existence_m1_subset_1
s: 148
d: [2,318,1,2,1,1,40,125,217]
a: [3,327,1,2,1,318]
s: 327
i: [1,116,0,4,1,0,0,255] t38_relat_1
a: [3,202,0,4,1,116]
s: 202
i: [1,115,0,4,1,0,0,255] t38_relat_1
a: [3,207,0,4,1,115]
s: 207
d: [2,329,1,2,1,1,40,126,197]
a: [3,334,1,2,1,329]
s: 334
i: [1,118,0,5,2,0,0,255] rc7_finset_1
s: 118
d: [2,330,1,2,1,1,40,126,217]
a: [3,339,1,2,1,330]
s: 339
i: [1,119,0,5,2,0,0,255] rc7_finset_1
s: 119
i: [1,121,0,5,2,0,0,3] rc2_finset_1
s: 121
i: [1,122,0,5,2,0,0,3] rc2_finset_1
s: 122
i: [1,127,0,5,2,0,0,1] t6_boole
s: 127
d: [2,346,1,4,2,1,44,310,127]
s: 346
i: [1,130,0,5,2,0,0,255] fc31_finset_1
s: 130
i: [1,131,0,5,2,0,0,3] fc17_finset_1
s: 131
d: [2,351,2,2,1,2,40,346,217]
a: [3,355,2,2,1,351]
s: 355
i: [1,134,0,6,3,0,0,255] cc9_finset_1
s: 134
d: [2,357,1,4,2,1,40,134,197]
d: [2,359,1,2,1,2,21,357,315]
a: [3,363,1,2,1,359]
s: 363
i: [1,117,0,7,2,0,0,255] rc7_finset_1
s: 117
i: [1,120,0,7,2,0,0,3] rc2_finset_1
s: 120
d: [2,347,1,4,2,1,44,315,127]
s: 347
i: [1,128,0,7,3,0,0,255] cc7_finset_1
s: 128
d: [2,374,2,4,2,1,44,363,127]
s: 374
r: 134
i: [1,135,0,7,3,0,0,255] fc36_finset_1
s: 135
i: [1,141,0,7,3,0,0,255] fc22_finset_1
s: 141
d: [2,345,1,5,2,1,44,305,127]
s: 345
i: [1,142,0,7,3,0,0,3] fc19_finset_1
s: 142
d: [2,358,1,4,2,1,40,134,217]
a: [3,368,1,2,1,358]
s: 368
i: [1,147,0,7,3,0,0,1] t8_boole
s: 147
i: [1,129,0,8,3,0,0,255] cc8_finset_1
s: 129
a: [3,371,1,2,1,358]
s: 371
i: [1,132,0,8,3,0,0,3] cc2_finset_1
s: 132
d: [2,378,1,5,2,0,40,128,148]
s: 378
i: [1,140,0,8,4,0,0,255] cc3_funcop_1
s: 140
d: [2,393,0,6,3,1,21,140,346]
d: [2,394,0,4,2,2,21,393,347]
s: 394
d: [2,395,1,2,1,3,40,394,197]
a: [3,400,1,2,1,395]
s: 400
i: [1,143,0,9,3,0,0,255] d18_relat_1
s: 143
d: [2,396,1,2,1,3,40,394,217]
a: [3,405,1,2,1,396]
s: 405
i: [1,144,0,9,3,0,0,255] d18_relat_1
s: 144
i: [1,150,0,11,3,0,0,2] redefinition_k1_relset_1
s: 150
d: [2,408,1,8,3,1,44,143,207]
d: [2,409,1,6,2,2,21,408,315]
s: 409
i: [1,151,0,11,3,0,0,2] dt_k1_relset_1
s: 151
d: [2,412,1,8,3,1,44,144,207]
d: [2,414,1,6,2,2,21,412,315]
s: 414
d: [2,421,2,0,0,3,52,305,149,414]
e: 421
d: [2,422,2,3,1,2,40,414,149]
d: [2,425,2,0,0,3,21,422,305]
e: 425
r: 345
r: 305
a: [3,304,0,3,1,299]
s: 304
a: [3,185,0,3,1,113]
d: [2,428,0,3,1,2,22,185,190]
a: [3,432,0,3,1,428]
s: 432
d: [2,437,1,5,2,1,44,304,127]
s: 437
i: [1,145,0,13,4,0,0,1] d2_partfun1
s: 145
d: [2,438,1,5,2,1,44,432,127]
s: 438
i: [1,146,0,13,4,0,0,1] d2_partfun1
s: 146
d: [2,343,1,6,2,1,44,202,127]
s: 343
d: [2,415,2,8,3,2,44,409,127]
s: 415
d: [2,344,1,6,2,1,44,207,127]
s: 344
d: [2,423,2,8,3,2,44,414,127]
s: 423
d: [2,387,1,6,2,0,40,129,148]
s: 387
d: [2,390,1,6,2,0,40,132,148]
s: 390
d: [2,455,2,10,4,1,44,144,344]
d: [2,458,2,8,3,2,21,455,347]
s: 458
d: [2,411,1,6,2,0,40,144,149]
s: 411
d: [2,462,3,8,3,2,40,411,415]
d: [2,467,3,6,2,2,21,462,347]
s: 467
d: [2,341,1,7,3,0,44,127,127]
s: 341
d: [2,388,1,7,3,0,40,129,117]
s: 388
d: [2,470,4,8,3,2,40,467,423]
d: [2,474,4,6,2,2,18,470]
s: 474
d: [2,389,1,7,3,0,40,129,120]
s: 389
d: [2,468,4,10,3,2,40,467,144]
s: 468
d: [2,419,1,15,5,0,44,151,150]
d: [2,420,1,10,3,0,18,419]
s: 420
d: [2,416,1,11,4,0,40,151,132]
s: 416
d: [2,499,2,9,3,1,44,420,207]
d: [2,501,2,7,2,2,21,499,315]
s: 501
d: [2,417,1,11,4,0,40,151,129]
s: 417
d: [2,500,2,11,4,1,44,420,344]
d: [2,502,2,9,3,2,21,500,347]
s: 502
d: [2,418,1,12,4,0,40,151,128]
s: 418
d: [2,442,1,17,6,0,44,150,145]
d: [2,443,1,12,4,0,18,442]
s: 443
d: [2,522,2,11,4,1,44,443,207]
d: [2,543,2,9,3,2,21,522,315]
s: 543
d: [2,441,1,17,6,0,44,151,145]
d: [2,444,1,12,4,0,18,441]
s: 444
d: [2,525,2,11,4,1,44,207,443]
d: [2,545,2,9,3,2,21,525,315]
s: 545
d: [2,440,1,17,6,0,44,145,150]
d: [2,445,1,12,4,0,18,440]
s: 445
d: [2,447,1,17,6,0,44,146,150]
d: [2,448,1,12,4,0,18,447]
s: 448
d: [2,552,2,0,0,4,52,315,304,432,207,448]
e: 552
i: [1,111,0,2,1,1,0,0] t2_pzfmisc1
i: [1,110,0,2,1,1,0,0] t2_pzfmisc1
f: 552,428,422,301,299,115,110,111,112,113
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pzfmisc1__t2_pzfmisc1
% SZS output start Proof for pzfmisc1__t2_pzfmisc1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (k1_xboole_0 = X0 => (v1_partfun1(X0,k1_xboole_0) & v1_funct_1(X0) & v4_relat_1(X0,k1_xboole_0) & v1_relat_1(X0)))) [input t2_pzfmisc1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (k1_xboole_0 = X0 => (v1_partfun1(X0,k1_xboole_0) & v1_funct_1(X0) & v4_relat_1(X0,k1_xboole_0) & v1_relat_1(X0)))) [negated conjecture 1]
33. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
34. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => k1_relset_1(X0,X1) = k9_xtuple_0(X1)) [input redefinition_k1_relset_1]
41. ! [X0] : ! [X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(X1,X0) <=> k1_relset_1(X0,X1) = X0)) [input d2_partfun1]
42. k1_xboole_0 = k10_xtuple_0(k1_xboole_0) & k1_xboole_0 = k9_xtuple_0(k1_xboole_0) [input t38_relat_1]
43. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0))) [input d18_relat_1]
44. ! [X0,X1] : (v1_relat_1(X1) => (v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0))) [flattening 43]
45. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(X1,X0) <=> k1_relset_1(X0,X1) = X0)) [flattening 41]
48. ! [X0] : r1_tarski(X0,X0) [rectify 33]
52. ? [X0] : (((~v1_partfun1(X0,k1_xboole_0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k1_xboole_0) | ~v1_relat_1(X0)) & k1_xboole_0 = X0) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
53. ? [X0] : ((~v1_partfun1(X0,k1_xboole_0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k1_xboole_0) | ~v1_relat_1(X0)) & k1_xboole_0 = X0 & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 52]
78. ! [X0,X1] : ((v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0)) | ~v1_relat_1(X1)) [ennf transformation 44]
79. ! [X0,X1] : ((v1_partfun1(X1,X0) <=> k1_relset_1(X0,X1) = X0) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 45]
80. ! [X0,X1] : ((v1_partfun1(X1,X0) <=> k1_relset_1(X0,X1) = X0) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 79]
82. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 34]
83. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 82]
86. ? [X0] : ((~v1_partfun1(X0,k1_xboole_0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k1_xboole_0) | ~v1_relat_1(X0)) & k1_xboole_0 = X0 & v1_funct_1(X0) & v1_relat_1(X0)) => ((~v1_partfun1(sK0,k1_xboole_0) | ~v1_funct_1(sK0) | ~v4_relat_1(sK0,k1_xboole_0) | ~v1_relat_1(sK0)) & k1_xboole_0 = sK0 & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
87. (~v1_partfun1(sK0,k1_xboole_0) | ~v1_funct_1(sK0) | ~v4_relat_1(sK0,k1_xboole_0) | ~v1_relat_1(sK0)) & k1_xboole_0 = sK0 & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 53,86]
92. ! [X0,X1] : (((v4_relat_1(X1,X0) | ~r1_tarski(k9_xtuple_0(X1),X0)) & (r1_tarski(k9_xtuple_0(X1),X0) | ~v4_relat_1(X1,X0))) | ~v1_relat_1(X1)) [nnf transformation 78]
93. ! [X0,X1] : (((v1_partfun1(X1,X0) | k1_relset_1(X0,X1) != X0) & (k1_relset_1(X0,X1) = X0 | ~v1_partfun1(X1,X0))) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 80]
110. v1_relat_1(sK0) [cnf transformation 87] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
111. v1_funct_1(sK0) [cnf transformation 87] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
112. k1_xboole_0 = sK0 [cnf transformation 87] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
113. ~v1_partfun1(sK0,k1_xboole_0) | ~v1_funct_1(sK0) | ~v4_relat_1(sK0,k1_xboole_0) | ~v1_relat_1(sK0) [cnf transformation 87] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
115. k1_xboole_0 = k9_xtuple_0(k1_xboole_0) [cnf transformation 42] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
144. ~r1_tarski(k9_xtuple_0(X1),X0) | v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 92] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
146. k1_relset_1(X0,X1) != X0 | v1_partfun1(X1,X0) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 93] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
149. r1_tarski(X0,X0) [cnf transformation 48] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
150. k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 83] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
171. 1 <=> v1_relat_1(sK0) [avatar definition]
172. v1_relat_1(sK0) <- (1) [avatar component clause 171] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
175. 2 <=> v4_relat_1(sK0,k1_xboole_0) [avatar definition]
177. ~v4_relat_1(sK0,k1_xboole_0) <- (~2) [avatar component clause 175] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
179. 3 <=> v1_funct_1(sK0) [avatar definition]
183. 4 <=> v1_partfun1(sK0,k1_xboole_0) [avatar definition]
185. ~v1_partfun1(sK0,k1_xboole_0) <- (~4) [avatar component clause 183] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
186. ~1 | ~2 | ~3 | ~4 [avatar split clause 113,183,179,175,171]
188. 5 <=> k1_xboole_0 = sK0 [avatar definition]
190. k1_xboole_0 = sK0 <- (5) [avatar component clause 188] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
191. 5 [avatar split clause 112,188]
192. 3 [avatar split clause 111,179]
193. 1 [avatar split clause 110,171]
205. 8 <=> k1_xboole_0 = k9_xtuple_0(k1_xboole_0) [avatar definition]
207. k1_xboole_0 = k9_xtuple_0(k1_xboole_0) <- (8) [avatar component clause 205] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
208. 8 [avatar split clause 115,205]
299. ~v4_relat_1(k1_xboole_0,k1_xboole_0) <- (~2, 5) [backward demodulation 177,190] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
301. v1_relat_1(k1_xboole_0) <- (1, 5) [backward demodulation 172,190] {a:0,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
303. 27 <=> v4_relat_1(k1_xboole_0,k1_xboole_0) [avatar definition]
304. v4_relat_1(k1_xboole_0,k1_xboole_0) <- (27) [avatar component clause 303] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
306. ~27 | 2 | ~5 [avatar split clause 299,188,175,303]
313. 29 <=> v1_relat_1(k1_xboole_0) [avatar definition]
315. v1_relat_1(k1_xboole_0) <- (29) [avatar component clause 313] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
316. 29 | ~1 | ~5 [avatar split clause 301,188,171,313]
412. ~r1_tarski(k1_xboole_0,X0) | v4_relat_1(k1_xboole_0,X0) | ~v1_relat_1(k1_xboole_0) <- (8) [superposition 144,207] {a:1,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
414. ~r1_tarski(k1_xboole_0,X0) | v4_relat_1(k1_xboole_0,X0) <- (8, 29) [subsumption resolution 412,315] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
422. v4_relat_1(k1_xboole_0,k1_xboole_0) <- (8, 29) [resolution 414,149] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
427. 27 | ~8 | ~29 [avatar split clause 422,313,205,303]
428. ~v1_partfun1(k1_xboole_0,k1_xboole_0) <- (~4, 5) [forward demodulation 185,190] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
430. 40 <=> v1_partfun1(k1_xboole_0,k1_xboole_0) [avatar definition]
432. ~v1_partfun1(k1_xboole_0,k1_xboole_0) <- (~40) [avatar component clause 430] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
433. ~40 | 4 | ~5 [avatar split clause 428,188,183,430]
447. k9_xtuple_0(X3) != X2 | v1_partfun1(X3,X2) | ~v4_relat_1(X3,X2) | ~v1_relat_1(X3) | ~v4_relat_1(X3,X2) | ~v1_relat_1(X3) [superposition 146,150] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
448. k9_xtuple_0(X3) != X2 | v1_partfun1(X3,X2) | ~v4_relat_1(X3,X2) | ~v1_relat_1(X3) [duplicate literal removal 447] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
552. $false <- (8, 27, 29, ~40) [unit resulting resolution 315,304,432,207,448] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
558. ~8 | ~27 | ~29 | 40 [avatar contradiction clause 552]
562. $false [avatar sat refutation 186,191,192,193,208,306,316,427,433,558] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:189,thDist:-189}
% SZS output end Proof for pzfmisc1__t2_pzfmisc1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 42
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 60
% Generated clauses: 292
% Active clauses: 100
% Passive clauses: 128
% Final active clauses: 94
% Final passive clauses: 21
% Main loop iterations started: 106
% 
% >>> Simplifying Inferences
% Duplicate literals: 40
% Fw subsumption resolutions: 23
% Fw demodulations: 7
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Forward subsumptions: 44
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 56
% Unit resulting resolution: 2
% Forward superposition: 39
% Backward superposition: 26
% Self superposition: 2
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 6
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 47
% SAT solver unit clauses: 24
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% non unit clause index maintenance: 0.001 s
% resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

