i: [1,50,0,2,1,1,0,0] t33_sysrel
a: [3,91,0,2,1,50]
s: 91
i: [1,51,0,16,2,1,0,0] t33_sysrel
a: [3,82,0,8,1,51]
s: 82
i: [1,52,0,3,1,0,0,1] dt_k4_relat_1
s: 52
i: [1,55,0,3,1,0,0,255] fc28_relat_1
s: 55
i: [1,70,0,3,1,0,0,1] reflexivity_r1_tarski
s: 70
i: [1,56,0,4,1,0,0,255] fc28_relat_1
s: 56
i: [1,57,0,4,1,0,0,255] fc28_relat_1
s: 57
i: [1,69,0,4,1,0,0,255] existence_m1_subset_1
s: 69
i: [1,76,0,4,1,0,0,255] rc3_relat_1
s: 76
i: [1,53,0,5,1,0,0,1] rd2_relat_1
s: 53
i: [1,54,0,5,1,0,0,1] rd1_relat_1
s: 54
i: [1,58,0,5,2,0,0,1] dt_k1_sysrel
s: 58
i: [1,71,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 71
i: [1,77,0,5,1,0,0,255] rc3_relat_1
s: 77
i: [1,78,0,5,1,0,0,255] rc3_relat_1
s: 78
i: [1,73,0,6,2,0,0,255] fc1_relat_1
s: 73
i: [1,67,0,7,2,0,0,255] t3_subset
s: 67
i: [1,68,0,7,2,0,0,255] t3_subset
s: 68
d: [2,94,1,5,1,0,40,67,69]
s: 94
i: [1,72,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 72
i: [1,61,0,8,2,0,0,1] t32_sysrel
s: 61
d: [2,96,1,6,2,0,44,73,72]
s: 96
i: [1,62,0,8,2,0,0,1] t32_sysrel
s: 62
d: [2,97,1,6,2,0,44,73,72]
s: 97
i: [1,63,0,8,2,0,0,1] t32_sysrel
s: 63
i: [1,64,0,8,2,0,0,1] t32_sysrel
s: 64
d: [2,98,1,9,2,0,44,61,54]
d: [2,99,1,6,1,0,21,98,52]
s: 99
i: [1,65,0,8,3,0,0,255] cc2_relat_1
s: 65
d: [2,103,1,9,2,0,44,62,53]
d: [2,104,1,6,1,0,21,103,52]
s: 104
i: [1,66,0,8,2,0,0,1] t15_sysrel
s: 66
i: [1,59,0,9,2,0,0,1] t26_sysrel
s: 59
d: [2,111,1,6,1,1,40,66,82]
a: [3,115,1,6,1,111]
s: 115
d: [2,124,2,0,0,2,52,91,115,61]
e: 124
d: [2,125,2,2,1,1,40,115,61]
d: [2,127,2,0,0,2,21,125,91]
e: 127
r: 82
r: 115
a: [3,81,0,8,1,51]
s: 81
a: [3,114,1,6,1,111]
s: 114
a: [3,86,0,8,1,51]
s: 86
d: [2,130,1,10,2,1,44,86,59]
d: [2,136,1,2,1,2,21,130,81]
d: [2,137,1,0,0,3,21,136,91]
e: 137
f: 137,127,111,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for sysrel__t33_sysrel
% SZS output start Proof for sysrel__t33_sysrel
1. ! [X0] : (v1_relat_1(X0) => (r1_tarski(k4_relat_1(k10_xtuple_0(k1_sysrel(X0))),k4_relat_1(k9_xtuple_0(X0))) & r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(X0))),k4_relat_1(k9_xtuple_0(X0))))) [input t33_sysrel]
2. ~! [X0] : (v1_relat_1(X0) => (r1_tarski(k4_relat_1(k10_xtuple_0(k1_sysrel(X0))),k4_relat_1(k9_xtuple_0(X0))) & r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(X0))),k4_relat_1(k9_xtuple_0(X0))))) [negated conjecture 1]
23. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k4_relat_1(X0),k4_relat_1(X1))) [input t15_sysrel]
24. ! [X0] : (v1_relat_1(X0) => k9_xtuple_0(k1_sysrel(X0)) = k10_xtuple_0(k1_sysrel(X0))) [input t26_sysrel]
25. ! [X0] : (v1_relat_1(X0) => (r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) & r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)) & r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) & r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)))) [input t32_sysrel]
27. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k4_relat_1(X0),k4_relat_1(X1))) [flattening 23]
31. ? [X0] : ((~r1_tarski(k4_relat_1(k10_xtuple_0(k1_sysrel(X0))),k4_relat_1(k9_xtuple_0(X0))) | ~r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(X0))),k4_relat_1(k9_xtuple_0(X0)))) & v1_relat_1(X0)) [ennf transformation 2]
33. ! [X0] : (k9_xtuple_0(k1_sysrel(X0)) = k10_xtuple_0(k1_sysrel(X0)) | ~v1_relat_1(X0)) [ennf transformation 24]
35. ! [X0] : ((r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) & r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)) & r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) & r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0))) | ~v1_relat_1(X0)) [ennf transformation 25]
37. ! [X0,X1] : (r1_tarski(k4_relat_1(X0),k4_relat_1(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 27]
43. ? [X0] : ((~r1_tarski(k4_relat_1(k10_xtuple_0(k1_sysrel(X0))),k4_relat_1(k9_xtuple_0(X0))) | ~r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(X0))),k4_relat_1(k9_xtuple_0(X0)))) & v1_relat_1(X0)) => ((~r1_tarski(k4_relat_1(k10_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) | ~r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0)))) & v1_relat_1(sK0)) [choice axiom]
44. (~r1_tarski(k4_relat_1(k10_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) | ~r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0)))) & v1_relat_1(sK0) [skolemisation 31,43]
50. v1_relat_1(sK0) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. ~r1_tarski(k4_relat_1(k10_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) | ~r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) [cnf transformation 44] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. k9_xtuple_0(k1_sysrel(X0)) = k10_xtuple_0(k1_sysrel(X0)) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
66. r1_tarski(k4_relat_1(X0),k4_relat_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 37] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
80. 1 <=> r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) [avatar definition]
81. r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) <- (1) [avatar component clause 80] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
82. ~r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) <- (~1) [avatar component clause 80] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. 2 <=> r1_tarski(k4_relat_1(k10_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) [avatar definition]
86. ~r1_tarski(k4_relat_1(k10_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) <- (~2) [avatar component clause 84] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. ~1 | ~2 [avatar split clause 51,84,80]
89. 3 <=> v1_relat_1(sK0) [avatar definition]
91. v1_relat_1(sK0) <- (3) [avatar component clause 89] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
92. 3 [avatar split clause 50,89]
111. ~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) <- (~1) [resolution 66,82] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
113. 4 <=> r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) [avatar definition]
115. ~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) <- (~4) [avatar component clause 113] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. ~4 | 1 [avatar split clause 111,80,113]
125. ~v1_relat_1(sK0) <- (~4) [resolution 115,61] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
127. $false <- (3, ~4) [subsumption resolution 125,91] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
128. ~3 | 4 [avatar contradiction clause 127]
130. ~r1_tarski(k4_relat_1(k9_xtuple_0(k1_sysrel(sK0))),k4_relat_1(k9_xtuple_0(sK0))) | ~v1_relat_1(sK0) <- (~2) [superposition 86,59] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
136. ~v1_relat_1(sK0) <- (1, ~2) [subsumption resolution 130,81] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
137. $false <- (1, ~2, 3) [subsumption resolution 136,91] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
138. ~1 | 2 | ~3 [avatar contradiction clause 137]
139. $false [avatar sat refutation 87,92,116,128,138] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:30,thDist:-30}
% SZS output end Proof for sysrel__t33_sysrel
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 68
% Active clauses: 33
% Passive clauses: 42
% Final active clauses: 31
% Final passive clauses: 7
% Main loop iterations started: 35
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 7
% Unit resulting resolution: 1
% Forward superposition: 9
% Backward superposition: 7
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 1
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit clause index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

