i: [1,42,0,2,1,1,0,0] t11_wellord2
a: [3,90,0,2,1,42]
s: 90
i: [1,41,0,2,1,1,0,0] t11_wellord2
a: [3,95,0,2,1,41]
s: 95
i: [1,40,0,2,1,1,0,0] t11_wellord2
a: [3,100,0,2,1,40]
s: 100
i: [1,45,0,3,1,1,0,0] t11_wellord2
a: [3,75,0,3,1,45]
s: 75
i: [1,44,0,4,1,1,0,0] t11_wellord2
a: [3,80,0,4,1,44]
s: 80
i: [1,43,0,4,1,1,0,0] t11_wellord2
a: [3,85,0,4,1,43]
s: 85
i: [1,71,0,2,1,0,0,1] rc1_ordinal1
a: [3,105,0,2,1,71]
s: 105
i: [1,46,0,3,1,0,0,1] dt_k1_wellord2
s: 46
i: [1,47,0,3,1,0,0,255] fc3_wellord2
s: 47
i: [1,48,0,3,1,0,0,255] fc3_wellord2
s: 48
i: [1,49,0,3,1,0,0,255] fc1_wellord2
s: 49
i: [1,50,0,3,1,0,0,255] fc1_wellord2
s: 50
i: [1,51,0,3,1,0,0,255] fc2_wellord2
s: 51
i: [1,52,0,3,1,0,0,255] fc2_wellord2
s: 52
i: [1,53,0,4,2,0,0,255] cc1_ordinal1
s: 53
i: [1,54,0,4,2,0,0,255] cc1_ordinal1
s: 54
d: [2,107,1,2,1,1,40,53,90]
a: [3,113,1,2,1,107]
s: 113
i: [1,55,0,5,2,0,0,255] fc4_wellord2
s: 55
d: [2,108,1,2,1,1,40,53,95]
a: [3,118,1,2,1,108]
s: 118
i: [1,56,0,5,2,0,0,255] fc4_wellord2
s: 56
i: [1,57,0,5,2,0,0,255] fc5_wellord2
s: 57
d: [2,125,1,2,1,1,40,54,90]
a: [3,131,1,2,1,125]
s: 131
i: [1,58,0,5,2,0,0,255] fc5_wellord2
s: 58
d: [2,126,1,2,1,1,40,54,95]
a: [3,136,1,2,1,126]
s: 136
i: [1,62,0,6,3,0,0,255] cc2_ordinal1
s: 62
i: [1,64,0,6,3,0,0,255] cc1_wellord1
s: 64
d: [2,109,1,2,1,1,40,53,105]
a: [3,123,1,2,1,109]
s: 123
i: [1,65,0,6,3,0,0,255] cc1_wellord1
s: 65
d: [2,127,1,2,1,1,40,54,105]
a: [3,141,1,2,1,127]
s: 141
i: [1,66,0,6,3,0,0,255] cc1_wellord1
s: 66
i: [1,67,0,6,3,0,0,255] cc1_wellord1
s: 67
i: [1,68,0,6,3,0,0,255] cc1_wellord1
s: 68
i: [1,60,0,10,4,0,0,1] t40_wellord1
s: 60
d: [2,146,1,9,3,1,40,60,80]
d: [2,148,1,6,2,1,21,146,46]
d: [2,149,1,4,1,2,21,148,100]
a: [3,153,1,4,1,149]
s: 153
i: [1,59,0,12,4,0,0,1] t10_wellord2
s: 59
i: [1,70,0,14,7,0,0,255] cc2_wellord1
s: 70
d: [2,147,1,9,3,1,40,60,85]
d: [2,155,1,6,2,1,21,147,46]
d: [2,156,1,4,1,2,21,155,100]
a: [3,160,1,4,1,156]
s: 160
i: [1,61,0,15,6,0,0,1] t42_wellord1
s: 61
d: [2,170,1,14,5,1,40,61,80]
d: [2,174,1,11,4,1,21,170,46]
d: [2,175,1,9,3,2,21,174,100]
s: 175
d: [2,171,1,14,5,1,40,61,85]
d: [2,176,1,11,4,1,21,171,46]
d: [2,177,1,9,3,2,21,176,100]
s: 177
d: [2,163,1,20,7,0,40,70,58]
d: [2,164,1,17,6,0,21,163,52]
d: [2,165,1,14,5,0,21,164,48]
d: [2,166,1,11,4,0,21,165,50]
d: [2,167,1,8,3,0,21,166,46]
d: [2,168,1,5,2,0,21,167,56]
s: 168
d: [2,172,2,14,5,1,40,61,153]
d: [2,178,2,12,4,2,21,172,100]
d: [2,179,2,9,3,2,21,178,46]
s: 179
d: [2,173,2,14,5,1,40,61,160]
d: [2,180,2,12,4,2,21,173,100]
d: [2,181,2,9,3,2,21,180,46]
s: 181
d: [2,205,3,5,2,3,40,179,80]
d: [2,208,3,3,1,3,21,205,100]
a: [3,212,3,3,1,208]
s: 212
d: [2,183,2,14,5,2,40,175,60]
d: [2,185,2,12,4,2,18,183]
d: [2,188,2,9,3,2,21,185,46]
s: 188
d: [2,184,2,14,5,2,40,175,59]
d: [2,189,2,11,4,2,21,184,46]
d: [2,190,2,9,3,3,21,189,90]
s: 190
d: [2,234,3,0,0,6,52,95,75,160,190]
e: 234
d: [2,236,3,5,2,4,40,190,160]
d: [2,239,3,2,1,5,21,236,75]
d: [2,240,3,0,0,6,21,239,95]
e: 240
f: 240,156,40,41,42,43,44,45
% Refutation found. Thanks to Tanya!
% SZS status Theorem for wellord2__t11_wellord2
% SZS output start Proof for wellord2__t11_wellord2
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => ((r4_wellord1(X0,k1_wellord2(X2)) & r4_wellord1(X0,k1_wellord2(X1))) => X1 = X2)))) [input t11_wellord2]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => ((r4_wellord1(X0,k1_wellord2(X2)) & r4_wellord1(X0,k1_wellord2(X1))) => X1 = X2)))) [negated conjecture 1]
7. ! [X0] : v1_relat_1(k1_wellord2(X0)) [input dt_k1_wellord2]
14. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r4_wellord1(k1_wellord2(X0),k1_wellord2(X1)) => X0 = X1))) [input t10_wellord2]
15. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r4_wellord1(X0,X1) => r4_wellord1(X1,X0)))) [input t40_wellord1]
16. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => ((r4_wellord1(X1,X2) & r4_wellord1(X0,X1)) => r4_wellord1(X0,X2))))) [input t42_wellord1]
17. ? [X0] : (? [X1] : (? [X2] : ((X1 != X2 & (r4_wellord1(X0,k1_wellord2(X2)) & r4_wellord1(X0,k1_wellord2(X1)))) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v1_relat_1(X0)) [ennf transformation 2]
18. ? [X0] : (? [X1] : (? [X2] : (X1 != X2 & r4_wellord1(X0,k1_wellord2(X2)) & r4_wellord1(X0,k1_wellord2(X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v1_relat_1(X0)) [flattening 17]
22. ! [X0] : (! [X1] : ((X0 = X1 | ~r4_wellord1(k1_wellord2(X0),k1_wellord2(X1))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 14]
23. ! [X0] : (! [X1] : (X0 = X1 | ~r4_wellord1(k1_wellord2(X0),k1_wellord2(X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 22]
24. ! [X0] : (! [X1] : ((r4_wellord1(X1,X0) | ~r4_wellord1(X0,X1)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 15]
25. ! [X0] : (! [X1] : (r4_wellord1(X1,X0) | ~r4_wellord1(X0,X1) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 24]
26. ! [X0] : (! [X1] : (! [X2] : ((r4_wellord1(X0,X2) | (~r4_wellord1(X1,X2) | ~r4_wellord1(X0,X1))) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 16]
27. ! [X0] : (! [X1] : (! [X2] : (r4_wellord1(X0,X2) | ~r4_wellord1(X1,X2) | ~r4_wellord1(X0,X1) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 26]
34. ? [X0] : (? [X1] : (? [X2] : (X1 != X2 & r4_wellord1(X0,k1_wellord2(X2)) & r4_wellord1(X0,k1_wellord2(X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (X1 != X2 & r4_wellord1(sK0,k1_wellord2(X2)) & r4_wellord1(sK0,k1_wellord2(X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v1_relat_1(sK0)) [choice axiom]
35. ? [X1] : (? [X2] : (X1 != X2 & r4_wellord1(sK0,k1_wellord2(X2)) & r4_wellord1(sK0,k1_wellord2(X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) => (? [X2] : (sK1 != X2 & r4_wellord1(sK0,k1_wellord2(X2)) & r4_wellord1(sK0,k1_wellord2(sK1)) & v3_ordinal1(X2)) & v3_ordinal1(sK1)) [choice axiom]
36. ? [X2] : (sK1 != X2 & r4_wellord1(sK0,k1_wellord2(X2)) & r4_wellord1(sK0,k1_wellord2(sK1)) & v3_ordinal1(X2)) => (sK1 != sK2 & r4_wellord1(sK0,k1_wellord2(sK2)) & r4_wellord1(sK0,k1_wellord2(sK1)) & v3_ordinal1(sK2)) [choice axiom]
37. ((sK1 != sK2 & r4_wellord1(sK0,k1_wellord2(sK2)) & r4_wellord1(sK0,k1_wellord2(sK1)) & v3_ordinal1(sK2)) & v3_ordinal1(sK1)) & v1_relat_1(sK0) [skolemisation 18,36,35,34]
40. v1_relat_1(sK0) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
41. v3_ordinal1(sK1) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
42. v3_ordinal1(sK2) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
43. r4_wellord1(sK0,k1_wellord2(sK1)) [cnf transformation 37] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
44. r4_wellord1(sK0,k1_wellord2(sK2)) [cnf transformation 37] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
45. sK1 != sK2 [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
46. v1_relat_1(k1_wellord2(X0)) [cnf transformation 7] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. ~r4_wellord1(k1_wellord2(X0),k1_wellord2(X1)) | X0 = X1 | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 23] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. ~r4_wellord1(X0,X1) | r4_wellord1(X1,X0) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 25] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. ~r4_wellord1(X1,X2) | r4_wellord1(X0,X2) | ~r4_wellord1(X0,X1) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. 1 <=> sK1 = sK2 [avatar definition]
75. sK1 != sK2 <- (~1) [avatar component clause 73] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
76. ~1 [avatar split clause 45,73]
78. 2 <=> r4_wellord1(sK0,k1_wellord2(sK2)) [avatar definition]
80. r4_wellord1(sK0,k1_wellord2(sK2)) <- (2) [avatar component clause 78] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
81. 2 [avatar split clause 44,78]
83. 3 <=> r4_wellord1(sK0,k1_wellord2(sK1)) [avatar definition]
85. r4_wellord1(sK0,k1_wellord2(sK1)) <- (3) [avatar component clause 83] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
86. 3 [avatar split clause 43,83]
88. 4 <=> v3_ordinal1(sK2) [avatar definition]
90. v3_ordinal1(sK2) <- (4) [avatar component clause 88] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
91. 4 [avatar split clause 42,88]
93. 5 <=> v3_ordinal1(sK1) [avatar definition]
95. v3_ordinal1(sK1) <- (5) [avatar component clause 93] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
96. 5 [avatar split clause 41,93]
98. 6 <=> v1_relat_1(sK0) [avatar definition]
100. v1_relat_1(sK0) <- (6) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
101. 6 [avatar split clause 40,98]
147. r4_wellord1(k1_wellord2(sK1),sK0) | ~v1_relat_1(k1_wellord2(sK1)) | ~v1_relat_1(sK0) <- (3) [resolution 60,85] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
155. r4_wellord1(k1_wellord2(sK1),sK0) | ~v1_relat_1(sK0) <- (3) [subsumption resolution 147,46] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
156. r4_wellord1(k1_wellord2(sK1),sK0) <- (3, 6) [subsumption resolution 155,100] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
158. 15 <=> r4_wellord1(k1_wellord2(sK1),sK0) [avatar definition]
160. r4_wellord1(k1_wellord2(sK1),sK0) <- (15) [avatar component clause 158] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
161. 15 | ~3 | ~6 [avatar split clause 156,98,83,158]
170. r4_wellord1(X0,k1_wellord2(sK2)) | ~r4_wellord1(X0,sK0) | ~v1_relat_1(k1_wellord2(sK2)) | ~v1_relat_1(sK0) | ~v1_relat_1(X0) <- (2) [resolution 61,80] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
174. r4_wellord1(X0,k1_wellord2(sK2)) | ~r4_wellord1(X0,sK0) | ~v1_relat_1(sK0) | ~v1_relat_1(X0) <- (2) [subsumption resolution 170,46] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
175. r4_wellord1(X0,k1_wellord2(sK2)) | ~r4_wellord1(X0,sK0) | ~v1_relat_1(X0) <- (2, 6) [subsumption resolution 174,100] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
184. ~r4_wellord1(k1_wellord2(X3),sK0) | ~v1_relat_1(k1_wellord2(X3)) | sK2 = X3 | ~v3_ordinal1(sK2) | ~v3_ordinal1(X3) <- (2, 6) [resolution 175,59] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
189. ~r4_wellord1(k1_wellord2(X3),sK0) | sK2 = X3 | ~v3_ordinal1(sK2) | ~v3_ordinal1(X3) <- (2, 6) [subsumption resolution 184,46] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
190. ~r4_wellord1(k1_wellord2(X3),sK0) | sK2 = X3 | ~v3_ordinal1(X3) <- (2, 4, 6) [subsumption resolution 189,90] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
236. sK1 = sK2 | ~v3_ordinal1(sK1) <- (2, 4, 6, 15) [resolution 190,160] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
239. ~v3_ordinal1(sK1) <- (~1, 2, 4, 6, 15) [subsumption resolution 236,75] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
240. $false <- (~1, 2, 4, 5, 6, 15) [subsumption resolution 239,95] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
241. 1 | ~2 | ~4 | ~5 | ~6 | ~15 [avatar contradiction clause 240]
242. $false [avatar sat refutation 76,81,86,91,96,101,161,241] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:145,thDist:-145}
% SZS output end Proof for wellord2__t11_wellord2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 32
% Generated clauses: 135
% Active clauses: 41
% Passive clauses: 52
% Final active clauses: 41
% Final passive clauses: 6
% Main loop iterations started: 46
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 32
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Equational tautologies: 2
% Forward subsumptions: 18
% 
% >>> Generating Inferences
% Binary resolution: 43
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 7
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

