i: [1,188,0,3,1,1,0,0] l12_xreal_0
a: [3,290,0,3,1,188]
s: 290
i: [1,189,0,2,1,0,0,255] fc1_numbers
a: [3,295,0,2,1,189]
s: 295
i: [1,190,0,2,1,0,0,1] spc1_boole
a: [3,300,0,2,1,190]
s: 300
i: [1,191,0,2,1,0,0,255] fc2_arytm_2
a: [3,305,0,2,1,191]
s: 305
i: [1,192,0,2,1,0,0,255] fc1_xboole_0
a: [3,310,0,2,1,192]
s: 310
i: [1,205,0,2,1,0,0,255] fc6_ordinal1
a: [3,370,0,2,1,205]
s: 370
i: [1,204,0,2,1,0,0,255] fc6_ordinal1
a: [3,375,0,2,1,204]
s: 375
i: [1,206,0,2,1,0,0,255] spc1_numerals
a: [3,395,0,2,1,206]
s: 395
i: [1,264,0,2,1,0,0,2] rc2_xboole_0
a: [3,407,0,2,1,264]
s: 407
i: [1,265,0,2,1,0,0,255] rc1_xreal_0
a: [3,412,0,2,1,265]
s: 412
i: [1,266,0,2,1,0,0,255] rc4_ordinal1
a: [3,417,0,2,1,266]
s: 417
i: [1,267,0,2,1,0,0,255] rc1_ordinal1
a: [3,422,0,2,1,267]
s: 422
i: [1,268,0,2,1,0,0,2] rc1_xxreal_0
a: [3,427,0,2,1,268]
s: 427
i: [1,269,0,2,1,0,0,2] rc1_xboole_0
a: [3,432,0,2,1,269]
s: 432
i: [1,273,0,2,1,0,0,255] rc2_ordinal1
a: [3,437,0,2,1,273]
s: 437
i: [1,272,0,2,1,0,0,255] rc2_ordinal1
a: [3,442,0,2,1,272]
s: 442
i: [1,271,0,2,1,0,0,255] rc2_ordinal1
a: [3,447,0,2,1,271]
s: 447
i: [1,270,0,2,1,0,0,255] rc2_ordinal1
a: [3,452,0,2,1,270]
s: 452
i: [1,275,0,2,1,0,0,255] rc5_ordinal1
a: [3,457,0,2,1,275]
s: 457
i: [1,274,0,2,1,0,0,255] rc5_ordinal1
a: [3,462,0,2,1,274]
s: 462
i: [1,276,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,467,0,2,1,276]
s: 467
i: [1,278,0,2,1,0,0,255] rc3_xxreal_0
a: [3,472,0,2,1,278]
s: 472
i: [1,277,0,2,1,0,0,255] rc3_xxreal_0
a: [3,477,0,2,1,277]
s: 477
i: [1,280,0,2,1,0,0,255] rc2_xxreal_0
a: [3,482,0,2,1,280]
s: 482
i: [1,279,0,2,1,0,0,255] rc2_xxreal_0
a: [3,487,0,2,1,279]
s: 487
i: [1,282,0,2,1,0,0,2] rc4_xxreal_0
a: [3,492,0,2,1,282]
s: 492
i: [1,281,0,2,1,0,0,2] rc4_xxreal_0
a: [3,497,0,2,1,281]
s: 497
i: [1,210,0,3,1,0,0,255] fc2_xboole_0
s: 210
i: [1,253,0,3,1,0,0,255] reflexivity_r1_tarski
s: 253
i: [1,193,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,315,0,3,1,193]
s: 315
i: [1,194,0,3,1,0,0,255] de_c2__xreal_0
a: [3,320,0,3,1,194]
s: 320
i: [1,195,0,3,1,0,0,255] de_c1__xreal_0
a: [3,325,0,3,1,195]
s: 325
i: [1,196,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,330,0,3,1,196]
s: 330
r: 310
d: [2,499,0,2,1,2,23,310,330]
a: [3,503,0,2,1,499]
s: 503
i: [1,197,0,3,1,0,0,255] l11_xreal_0
a: [3,335,0,3,1,197]
s: 335
i: [1,198,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,340,0,3,1,198]
s: 340
i: [1,199,0,3,1,0,0,255] dt_c2__xreal_0
a: [3,345,0,3,1,199]
s: 345
i: [1,200,0,3,1,0,0,255] dt_c1__xreal_0
a: [3,350,0,3,1,200]
s: 350
i: [1,201,0,3,1,0,0,255] t1_numerals
a: [3,355,0,3,1,201]
s: 355
i: [1,209,0,3,1,0,0,255] spc1_numerals
a: [3,380,0,3,1,209]
s: 380
i: [1,208,0,3,1,0,0,255] spc1_numerals
a: [3,385,0,3,1,208]
s: 385
d: [2,505,0,3,1,2,22,335,320]
d: [2,506,0,3,1,3,22,505,325]
a: [3,510,0,3,1,506]
s: 510
d: [2,512,0,3,1,2,22,345,320]
a: [3,516,0,3,1,512]
s: 516
d: [2,518,0,3,1,2,22,350,325]
a: [3,522,0,3,1,518]
s: 522
d: [2,524,0,3,1,2,22,355,330]
a: [3,528,0,3,1,524]
s: 528
d: [2,530,0,3,1,2,22,385,315]
a: [3,534,0,3,1,530]
s: 534
i: [1,212,0,4,2,0,0,255] cc4_xreal_0
s: 212
i: [1,216,0,4,2,0,0,255] cc2_xreal_0
s: 216
d: [2,536,1,2,1,1,40,212,412]
a: [3,540,1,2,1,536]
s: 540
i: [1,217,0,4,2,0,0,255] cc6_ordinal1
s: 217
i: [1,218,0,4,2,0,0,255] cc2_xxreal_0
s: 218
d: [2,542,1,2,1,1,40,216,417]
a: [3,547,1,2,1,542]
s: 547
i: [1,219,0,4,2,0,0,255] cc1_ordinal1
s: 219
d: [2,543,1,2,1,1,40,216,457]
a: [3,552,1,2,1,543]
s: 552
i: [1,220,0,4,2,0,0,255] cc1_ordinal1
s: 220
i: [1,222,0,4,2,0,0,255] cc9_ordinal1
s: 222
d: [2,554,1,2,1,1,40,217,417]
a: [3,559,1,2,1,554]
s: 559
i: [1,223,0,4,2,0,0,255] cc7_ordinal1
s: 223
d: [2,555,1,2,1,1,40,217,457]
a: [3,564,1,2,1,555]
s: 564
i: [1,224,0,4,2,0,0,255] cc3_ordinal1
s: 224
i: [1,252,0,4,1,0,0,255] existence_m1_subset_1
s: 252
d: [2,566,1,2,1,1,40,218,417]
a: [3,571,1,2,1,566]
s: 571
i: [1,254,0,4,1,0,0,255] fc3_xboole_0
s: 254
d: [2,567,1,2,1,1,40,218,457]
a: [3,576,1,2,1,567]
s: 576
i: [1,202,0,4,1,0,0,255] dt_k5_numbers
a: [3,360,0,4,1,202]
s: 360
i: [1,203,0,4,1,0,0,255] dt_k6_numbers
a: [3,365,0,4,1,203]
s: 365
d: [2,580,1,2,1,1,40,219,370]
a: [3,586,1,2,1,580]
s: 586
i: [1,207,0,4,1,0,0,255] spc1_numerals
a: [3,390,0,4,1,207]
s: 390
d: [2,581,1,2,1,1,40,219,422]
a: [3,591,1,2,1,581]
s: 591
d: [2,686,0,4,1,2,22,360,315]
a: [3,690,0,4,1,686]
s: 690
d: [2,692,0,4,1,2,22,365,315]
a: [3,696,0,4,1,692]
s: 696
d: [2,595,1,2,1,1,40,220,370]
a: [3,601,1,2,1,595]
s: 601
d: [2,698,0,4,1,2,22,390,315]
a: [3,702,0,4,1,698]
s: 702
d: [2,596,1,2,1,1,40,220,422]
a: [3,606,1,2,1,596]
s: 606
i: [1,225,0,5,2,0,0,255] t6_boole
s: 225
i: [1,226,0,5,2,0,0,255] cc8_ordinal1
s: 226
d: [2,608,1,2,1,1,40,222,497]
a: [3,614,1,2,1,608]
s: 614
i: [1,227,0,5,2,0,0,255] cc1_xreal_0
s: 227
d: [2,609,1,2,1,1,40,222,432]
a: [3,619,1,2,1,609]
s: 619
i: [1,249,0,5,2,0,0,255] t7_boole
s: 249
i: [1,261,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,402,0,5,2,261]
s: 402
d: [2,610,1,2,1,1,40,222,503]
a: [3,624,1,2,1,610]
s: 624
d: [2,704,0,5,2,1,22,225,330]
s: 704
d: [2,732,1,5,2,2,44,290,704]
s: 732
i: [1,232,0,6,3,0,0,255] cc2_ordinal1
s: 232
i: [1,235,0,6,3,0,0,255] cc5_xxreal_0
s: 235
d: [2,638,1,2,1,1,40,223,497]
a: [3,644,1,2,1,638]
s: 644
i: [1,237,0,6,3,0,0,255] cc5_xxreal_0
s: 237
d: [2,639,1,2,1,1,40,223,432]
a: [3,649,1,2,1,639]
s: 649
i: [1,238,0,6,3,0,0,255] cc3_xxreal_0
s: 238
i: [1,240,0,6,3,0,0,255] cc3_xxreal_0
s: 240
d: [2,640,1,2,1,1,40,223,503]
a: [3,654,1,2,1,640]
s: 654
i: [1,242,0,6,3,0,0,255] cc7_xxreal_0
s: 242
d: [2,668,1,2,1,1,40,224,497]
a: [3,674,1,2,1,668]
s: 674
i: [1,243,0,6,3,0,0,255] cc7_xxreal_0
s: 243
i: [1,245,0,6,2,0,0,255] t1_subset
s: 245
d: [2,669,1,2,1,1,40,224,432]
a: [3,679,1,2,1,669]
s: 679
i: [1,256,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 256
d: [2,670,1,2,1,1,40,224,503]
a: [3,684,1,2,1,670]
s: 684
i: [1,221,0,7,3,0,0,255] cc5_ordinal1
s: 221
i: [1,247,0,7,2,0,0,255] t3_subset
s: 247
d: [2,706,1,2,1,1,40,226,534]
a: [3,712,1,2,1,706]
s: 712
i: [1,248,0,7,3,0,0,255] t8_boole
s: 248
d: [2,719,1,2,1,1,40,227,380]
a: [3,724,1,2,1,719]
s: 724
i: [1,255,0,7,1,0,0,255] commutativity_k2_tarski
s: 255
i: [1,211,0,8,3,0,0,255] cc10_ordinal1
s: 211
d: [2,744,1,4,2,1,40,235,472]
d: [2,745,1,2,1,2,21,744,477]
a: [3,749,1,2,1,745]
s: 749
i: [1,229,0,8,4,0,0,255] cc4_xxreal_0
s: 229
d: [2,760,1,4,2,1,40,237,395]
a: [3,764,1,2,1,760]
s: 764
i: [1,231,0,8,4,0,0,255] cc6_xxreal_0
s: 231
i: [1,233,0,8,4,0,0,255] cc8_xxreal_0
s: 233
a: [3,769,1,2,1,760]
s: 769
i: [1,246,0,8,3,0,0,255] t2_subset
s: 246
d: [2,761,1,4,2,1,40,237,482]
d: [2,771,1,2,1,2,21,761,487]
a: [3,775,1,2,1,771]
s: 775
i: [1,251,0,9,3,0,0,255] t5_subset
s: 251
i: [1,244,0,10,1,0,0,255] d5_tarski
s: 244
d: [2,792,1,4,2,1,40,238,482]
d: [2,793,1,2,1,2,21,792,487]
a: [3,797,1,2,1,793]
s: 797
i: [1,250,0,10,3,0,0,255] t4_subset
s: 250
d: [2,853,1,4,2,1,40,221,380]
a: [3,865,1,2,1,853]
s: 865
i: [1,262,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 262
i: [1,263,0,12,4,0,0,255] connectedness_r1_arytm_2
s: 263
d: [2,1027,1,7,3,1,40,262,290]
d: [2,1031,1,5,2,2,21,1027,764]
a: [3,1035,1,3,1,1031]
s: 1035
i: [1,260,0,14,4,0,0,255] existence_m2_subset_1
s: 260
d: [2,1041,2,5,2,2,44,1035,704]
s: 1041
i: [1,257,0,15,5,0,0,255] dt_m2_subset_1
s: 257
i: [1,258,0,15,5,0,0,255] redefinition_m2_subset_1
s: 258
d: [2,1028,2,9,4,2,40,262,732]
d: [2,1037,2,7,3,3,21,1028,764]
s: 1037
i: [1,259,0,15,5,0,0,255] redefinition_m2_subset_1
s: 259
d: [2,886,1,4,2,1,40,211,690]
a: [3,891,1,2,1,886]
s: 891
i: [1,215,0,16,5,0,0,255] l3_xreal_0
s: 215
i: [1,286,0,22,8,0,0,255] l3_xreal_0
s: 286
d: [2,899,1,6,3,1,40,229,427]
a: [3,933,1,2,1,899]
s: 933
r: 933
d: [2,1066,0,19,7,0,21,286,245]
d: [2,1067,0,16,6,0,21,1066,245]
s: 1067
d: [2,705,1,2,1,1,40,226,528]
d: [2,801,2,2,1,1,40,654,216]
d: [2,955,1,5,2,1,40,246,522]
d: [2,962,1,3,1,2,21,955,305]
d: [2,958,1,5,2,1,40,246,516]
d: [2,980,1,3,1,2,21,958,305]
d: [2,1086,1,13,5,1,40,1067,510]
d: [2,1087,1,10,4,2,21,1086,290]
d: [2,1088,1,8,3,3,21,1087,724]
f: 1088,980,962,801,719,705,524,518,512,506,209,201,200,199,197,196,195,194,191,188
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_0__l12_xreal_0
% SZS output start Proof for xreal_0__l12_xreal_0
1. r1_xxreal_0(k6_numbers,1) [input l12_xreal_0]
2. ~r1_xxreal_0(k6_numbers,1) [negated conjecture 1]
11. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
17. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
24. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
38. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xreal_0(X0)) [input cc1_xreal_0]
57. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
65. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
72. m1_subset_1(c1__xreal_0,k2_arytm_2) [input dt_c1__xreal_0]
73. m1_subset_1(c2__xreal_0,k2_arytm_2) [input dt_c2__xreal_0]
74. k6_numbers = c1__xreal_0 [input de_c1__xreal_0]
75. 1 = c2__xreal_0 [input de_c2__xreal_0]
78. ~v1_xboole_0(k2_arytm_2) [input fc2_arytm_2]
82. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
85. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
87. r1_arytm_2(c1__xreal_0,c2__xreal_0) [input l11_xreal_0]
88. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (~(~(r2_hidden(X0,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2)) & r2_hidden(X1,k2_arytm_2)) & ~(? [X2] : (? [X3] : (r1_arytm_2(X3,X2) & k4_tarski(k6_numbers,X3) = X1 & k4_tarski(k6_numbers,X2) = X0 & m1_subset_1(X3,k2_arytm_2)) & m1_subset_1(X2,k2_arytm_2)) & r2_hidden(X1,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2)) & r2_hidden(X0,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2))) & ~(? [X2] : (? [X3] : (r1_arytm_2(X2,X3) & X1 = X3 & X0 = X2 & m1_subset_1(X3,k2_arytm_2)) & m1_subset_1(X2,k2_arytm_2)) & r2_hidden(X1,k2_arytm_2) & r2_hidden(X0,k2_arytm_2))) => r1_xxreal_0(X0,X1)))) [input l3_xreal_0]
89. ~r1_xxreal_0(k6_numbers,1) [flattening 2]
90. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (~(~(r2_hidden(X0,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2)) & r2_hidden(X1,k2_arytm_2)) & ~(? [X2] : (? [X3] : (r1_arytm_2(X3,X2) & k4_tarski(k6_numbers,X3) = X1 & k4_tarski(k6_numbers,X2) = X0 & m1_subset_1(X3,k2_arytm_2)) & m1_subset_1(X2,k2_arytm_2)) & r2_hidden(X1,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2)) & r2_hidden(X0,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2))) & ~(? [X4] : (? [X5] : (r1_arytm_2(X4,X5) & X1 = X5 & X0 = X4 & m1_subset_1(X5,k2_arytm_2)) & m1_subset_1(X4,k2_arytm_2)) & r2_hidden(X1,k2_arytm_2) & r2_hidden(X0,k2_arytm_2))) => r1_xxreal_0(X0,X1)))) [rectify 88]
92. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 82]
93. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 57]
110. ! [X0] : (! [X1] : ((r1_xxreal_0(X0,X1) | ((~r2_hidden(X0,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2)) | ~r2_hidden(X1,k2_arytm_2)) & (! [X2] : (! [X3] : (~r1_arytm_2(X3,X2) | k4_tarski(k6_numbers,X3) != X1 | k4_tarski(k6_numbers,X2) != X0 | ~m1_subset_1(X3,k2_arytm_2)) | ~m1_subset_1(X2,k2_arytm_2)) | ~r2_hidden(X1,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2)) | ~r2_hidden(X0,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2))) & (! [X4] : (! [X5] : (~r1_arytm_2(X4,X5) | X1 != X5 | X0 != X4 | ~m1_subset_1(X5,k2_arytm_2)) | ~m1_subset_1(X4,k2_arytm_2)) | ~r2_hidden(X1,k2_arytm_2) | ~r2_hidden(X0,k2_arytm_2)))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 90]
111. ! [X0] : (! [X1] : (r1_xxreal_0(X0,X1) | ((~r2_hidden(X0,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2)) | ~r2_hidden(X1,k2_arytm_2)) & (! [X2] : (! [X3] : (~r1_arytm_2(X3,X2) | k4_tarski(k6_numbers,X3) != X1 | k4_tarski(k6_numbers,X2) != X0 | ~m1_subset_1(X3,k2_arytm_2)) | ~m1_subset_1(X2,k2_arytm_2)) | ~r2_hidden(X1,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2)) | ~r2_hidden(X0,k2_zfmisc_1(k1_tarski(k6_numbers),k2_arytm_2))) & (! [X4] : (! [X5] : (~r1_arytm_2(X4,X5) | X1 != X5 | X0 != X4 | ~m1_subset_1(X5,k2_arytm_2)) | ~m1_subset_1(X4,k2_arytm_2)) | ~r2_hidden(X1,k2_arytm_2) | ~r2_hidden(X0,k2_arytm_2))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 110]
112. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
121. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 17]
122. ! [X0] : (v1_xreal_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 38]
137. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 92]
138. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 93]
139. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 138]
188. ~r1_xxreal_0(k6_numbers,1) [cnf transformation 89] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
191. ~v1_xboole_0(k2_arytm_2) [cnf transformation 78] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
194. 1 = c2__xreal_0 [cnf transformation 75] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
195. k6_numbers = c1__xreal_0 [cnf transformation 74] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
196. k6_numbers = k1_xboole_0 [cnf transformation 65] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
197. r1_arytm_2(c1__xreal_0,c2__xreal_0) [cnf transformation 87] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
199. m1_subset_1(c2__xreal_0,k2_arytm_2) [cnf transformation 73] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
200. m1_subset_1(c1__xreal_0,k2_arytm_2) [cnf transformation 72] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
201. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 24] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
209. m1_subset_1(1,k1_numbers) [cnf transformation 85] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
213. r1_xxreal_0(X0,X1) | ~r1_arytm_2(X4,X5) | X1 != X5 | X0 != X4 | ~m1_subset_1(X5,k2_arytm_2) | ~m1_subset_1(X4,k2_arytm_2) | ~r2_hidden(X1,k2_arytm_2) | ~r2_hidden(X0,k2_arytm_2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 111] {a:0,w:28,wCS:11200,sine:255,thAx:0,allAx:1,thDist:-1}
216. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 112] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
226. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 121] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
227. ~m1_subset_1(X0,k1_numbers) | v1_xreal_0(X0) [cnf transformation 122] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
245. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 137] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
246. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 139] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
285. r1_xxreal_0(X0,X5) | ~r1_arytm_2(X4,X5) | X0 != X4 | ~m1_subset_1(X5,k2_arytm_2) | ~m1_subset_1(X4,k2_arytm_2) | ~r2_hidden(X5,k2_arytm_2) | ~r2_hidden(X0,k2_arytm_2) | ~v1_xreal_0(X5) | ~v1_xreal_0(X0) [equality resolution 213] {a:0,w:25,wCS:10000,sine:255,thAx:0,allAx:1,thDist:-1}
286. r1_xxreal_0(X4,X5) | ~r1_arytm_2(X4,X5) | ~m1_subset_1(X5,k2_arytm_2) | ~m1_subset_1(X4,k2_arytm_2) | ~r2_hidden(X5,k2_arytm_2) | ~r2_hidden(X4,k2_arytm_2) | ~v1_xreal_0(X5) | ~v1_xreal_0(X4) [equality resolution 285] {a:0,w:22,wCS:8800,sine:255,thAx:0,allAx:1,thDist:-1}
288. 1 <=> r1_xxreal_0(k6_numbers,1) [avatar definition]
290. ~r1_xxreal_0(k6_numbers,1) <- (~1) [avatar component clause 288] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
291. ~1 [avatar split clause 188,288]
303. 4 <=> v1_xboole_0(k2_arytm_2) [avatar definition]
305. ~v1_xboole_0(k2_arytm_2) <- (~4) [avatar component clause 303] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
306. ~4 [avatar split clause 191,303]
318. 7 <=> 1 = c2__xreal_0 [avatar definition]
320. 1 = c2__xreal_0 <- (7) [avatar component clause 318] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
321. 7 [avatar split clause 194,318]
323. 8 <=> k6_numbers = c1__xreal_0 [avatar definition]
325. k6_numbers = c1__xreal_0 <- (8) [avatar component clause 323] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
326. 8 [avatar split clause 195,323]
328. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
330. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 328] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
331. 9 [avatar split clause 196,328]
333. 10 <=> r1_arytm_2(c1__xreal_0,c2__xreal_0) [avatar definition]
335. r1_arytm_2(c1__xreal_0,c2__xreal_0) <- (10) [avatar component clause 333] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
336. 10 [avatar split clause 197,333]
343. 12 <=> m1_subset_1(c2__xreal_0,k2_arytm_2) [avatar definition]
345. m1_subset_1(c2__xreal_0,k2_arytm_2) <- (12) [avatar component clause 343] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
346. 12 [avatar split clause 199,343]
348. 13 <=> m1_subset_1(c1__xreal_0,k2_arytm_2) [avatar definition]
350. m1_subset_1(c1__xreal_0,k2_arytm_2) <- (13) [avatar component clause 348] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
351. 13 [avatar split clause 200,348]
353. 14 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
355. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (14) [avatar component clause 353] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
356. 14 [avatar split clause 201,353]
378. 19 <=> m1_subset_1(1,k1_numbers) [avatar definition]
380. m1_subset_1(1,k1_numbers) <- (19) [avatar component clause 378] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
381. 19 [avatar split clause 209,378]
505. r1_arytm_2(c1__xreal_0,1) <- (7, 10) [forward demodulation 335,320] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
506. r1_arytm_2(k6_numbers,1) <- (7, 8, 10) [forward demodulation 505,325] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
508. 45 <=> r1_arytm_2(k6_numbers,1) [avatar definition]
510. r1_arytm_2(k6_numbers,1) <- (45) [avatar component clause 508] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
511. 45 | ~7 | ~8 | ~10 [avatar split clause 506,333,323,318,508]
512. m1_subset_1(1,k2_arytm_2) <- (7, 12) [forward demodulation 345,320] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
514. 46 <=> m1_subset_1(1,k2_arytm_2) [avatar definition]
516. m1_subset_1(1,k2_arytm_2) <- (46) [avatar component clause 514] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
517. 46 | ~7 | ~12 [avatar split clause 512,343,318,514]
518. m1_subset_1(k6_numbers,k2_arytm_2) <- (8, 13) [forward demodulation 350,325] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
520. 47 <=> m1_subset_1(k6_numbers,k2_arytm_2) [avatar definition]
522. m1_subset_1(k6_numbers,k2_arytm_2) <- (47) [avatar component clause 520] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
523. 47 | ~8 | ~13 [avatar split clause 518,348,323,520]
524. m1_subset_1(k6_numbers,k4_ordinal1) <- (9, 14) [forward demodulation 355,330] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
526. 48 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
528. m1_subset_1(k6_numbers,k4_ordinal1) <- (48) [avatar component clause 526] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
529. 48 | ~9 | ~14 [avatar split clause 524,353,328,526]
652. 68 <=> v7_ordinal1(k6_numbers) [avatar definition]
654. v7_ordinal1(k6_numbers) <- (68) [avatar component clause 652] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
705. v7_ordinal1(k6_numbers) <- (48) [resolution 226,528] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
708. 68 | ~48 [avatar split clause 705,526,652]
719. v1_xreal_0(1) <- (19) [resolution 227,380] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
722. 79 <=> v1_xreal_0(1) [avatar definition]
724. v1_xreal_0(1) <- (79) [avatar component clause 722] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
725. 79 | ~19 [avatar split clause 719,378,722]
801. v1_xreal_0(k6_numbers) <- (68) [resolution 654,216] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
810. 90 <=> v1_xreal_0(k6_numbers) [avatar definition]
813. 90 | ~68 [avatar split clause 801,652,810]
955. v1_xboole_0(k2_arytm_2) | r2_hidden(k6_numbers,k2_arytm_2) <- (47) [resolution 246,522] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:2,thDist:-2}
958. v1_xboole_0(k2_arytm_2) | r2_hidden(1,k2_arytm_2) <- (46) [resolution 246,516] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:2,thDist:-2}
962. r2_hidden(k6_numbers,k2_arytm_2) <- (~4, 47) [subsumption resolution 955,305] {a:1,w:3,wCS:1200,sine:255,thAx:0,allAx:2,thDist:-2}
964. 113 <=> r2_hidden(k6_numbers,k2_arytm_2) [avatar definition]
967. 113 | 4 | ~47 [avatar split clause 962,520,303,964]
980. r2_hidden(1,k2_arytm_2) <- (~4, 46) [subsumption resolution 958,305] {a:1,w:3,wCS:1200,sine:255,thAx:0,allAx:2,thDist:-2}
982. 116 <=> r2_hidden(1,k2_arytm_2) [avatar definition]
985. 116 | 4 | ~46 [avatar split clause 980,514,303,982]
1066. r1_xxreal_0(X4,X5) | ~r1_arytm_2(X4,X5) | ~m1_subset_1(X4,k2_arytm_2) | ~r2_hidden(X5,k2_arytm_2) | ~r2_hidden(X4,k2_arytm_2) | ~v1_xreal_0(X5) | ~v1_xreal_0(X4) [subsumption resolution 286,245] {a:0,w:19,wCS:7600,sine:255,thAx:0,allAx:1,thDist:-1}
1067. ~r1_arytm_2(X4,X5) | r1_xxreal_0(X4,X5) | ~r2_hidden(X5,k2_arytm_2) | ~r2_hidden(X4,k2_arytm_2) | ~v1_xreal_0(X5) | ~v1_xreal_0(X4) [subsumption resolution 1066,245] {a:0,w:16,wCS:6400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1086. r1_xxreal_0(k6_numbers,1) | ~r2_hidden(1,k2_arytm_2) | ~r2_hidden(k6_numbers,k2_arytm_2) | ~v1_xreal_0(1) | ~v1_xreal_0(k6_numbers) <- (45) [resolution 1067,510] {a:1,w:13,wCS:5200,sine:255,thAx:0,allAx:2,thDist:-2}
1087. ~r2_hidden(1,k2_arytm_2) | ~r2_hidden(k6_numbers,k2_arytm_2) | ~v1_xreal_0(1) | ~v1_xreal_0(k6_numbers) <- (~1, 45) [subsumption resolution 1086,290] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1088. ~r2_hidden(1,k2_arytm_2) | ~r2_hidden(k6_numbers,k2_arytm_2) | ~v1_xreal_0(k6_numbers) <- (~1, 45, 79) [subsumption resolution 1087,724] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1089. ~90 | ~113 | ~116 | 1 | ~45 | ~79 [avatar split clause 1088,722,508,288,982,964,810]
1090. $false [avatar sat refutation 291,306,321,326,331,336,346,351,356,381,511,517,523,529,708,725,813,967,985,1089] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:544,thDist:-544}
% SZS output end Proof for xreal_0__l12_xreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 87
% >>> Preprocessing
% Introduced skolems: 14
% Pure predicates: 3
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 95
% Generated clauses: 384
% Active clauses: 117
% Passive clauses: 199
% Final active clauses: 115
% Final passive clauses: 65
% Main loop iterations started: 133
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Fw subsumption resolutions: 21
% Fw demodulations: 10
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 33
% 
% >>> Generating Inferences
% Binary resolution: 116
% Factoring: 1
% Forward superposition: 8
% Backward superposition: 12
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 28
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 128
% SAT solver unit clauses: 43
% SAT solver binary clauses: 53
% 
% Memory used [KB]: 5500
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward demodulation index maintenance: 0.001 s
% other: 0.005 s ( own 0.003 s ) 

