i: [1,31,0,2,1,1,0,0] t120_group_9
a: [3,57,0,2,1,31]
s: 57
i: [1,30,0,2,1,1,0,0] t120_group_9
a: [3,62,0,2,1,30]
s: 62
i: [1,33,0,15,2,1,0,0] t120_group_9
a: [3,47,0,6,1,33]
s: 47
a: [3,50,0,9,1,33]
s: 50
i: [1,34,0,5,2,0,0,1] dt_k2_relat_1
s: 34
i: [1,40,0,6,2,0,0,1] dt_k5_relat_1
s: 40
d: [2,64,1,6,2,1,44,40,50]
a: [3,68,1,3,1,64]
s: 68
d: [2,74,2,0,0,2,52,57,68,34]
e: 74
d: [2,75,2,2,1,1,40,68,34]
d: [2,77,2,0,0,2,21,75,57]
e: 77
r: 68
i: [1,41,0,6,2,0,0,1] dt_k6_relat_1
s: 41
a: [3,67,1,3,1,64]
s: 67
i: [1,35,0,7,2,0,0,1] involutiveness_k2_relat_1
s: 35
i: [1,36,0,8,2,0,0,1] rd4_relat_1
s: 36
a: [3,72,1,3,1,64]
s: 72
i: [1,37,0,8,2,0,0,1] rd6_relat_1
s: 37
i: [1,42,0,11,2,0,0,1] rd7_relat_1
s: 42
i: [1,43,0,11,2,0,0,1] rd5_relat_1
s: 43
r: 47
r: 50
a: [3,46,0,6,1,33]
s: 46
d: [2,99,1,8,2,1,44,42,46]
d: [2,101,1,6,1,2,21,99,57]
a: [3,105,1,6,1,101]
s: 105
a: [3,51,0,9,1,33]
s: 51
i: [1,38,0,19,4,0,0,1] l168_group_9
s: 38
d: [2,109,1,0,0,4,52,62,57,46,51,38]
e: 109
d: [2,111,1,16,4,1,44,38,46]
d: [2,116,1,13,3,1,19,111]
d: [2,120,1,4,2,2,21,116,51]
d: [2,121,1,2,1,3,21,120,57]
d: [2,122,1,0,0,4,21,121,62]
e: 122
r: 46
r: 51
s: 47
s: 50
d: [2,88,1,12,2,1,44,43,50]
d: [2,93,1,9,1,2,21,88,67]
a: [3,97,1,9,1,93]
s: 97
i: [1,39,0,19,4,0,0,1] l168_group_9
s: 39
d: [2,133,1,0,0,4,52,62,57,47,50,39]
e: 133
d: [2,136,1,15,4,1,44,39,50]
d: [2,143,1,10,3,1,19,136]
d: [2,147,1,4,2,2,21,143,47]
d: [2,148,1,2,1,3,21,147,57]
d: [2,149,1,0,0,4,21,148,62]
e: 149
i: [1,32,0,15,2,1,0,0] t120_group_9
f: 149,122,30,31,32,33
% Refutation found. Thanks to Tanya!
% SZS status Theorem for group_9__t120_group_9
% SZS output start Proof for group_9__t120_group_9
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (k6_relat_1(k10_xtuple_0(X0),X1) = X0 <=> k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0)))))) [input t120_group_9]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (k6_relat_1(k10_xtuple_0(X0),X1) = X0 <=> k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0)))))) [negated conjecture 1]
13. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (k6_relat_1(k10_xtuple_0(X0),X1) = X0 <=> k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0)))))) [input l168_group_9]
14. ? [X0] : (? [X1] : ((k6_relat_1(k10_xtuple_0(X0),X1) = X0 <~> k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0)))) & v1_relat_1(X1)) & v1_relat_1(X0)) [ennf transformation 2]
19. ! [X0] : (! [X1] : ((k6_relat_1(k10_xtuple_0(X0),X1) = X0 <=> k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0)))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 13]
24. ? [X0] : (? [X1] : (((k2_relat_1(X0) != k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0))) | k6_relat_1(k10_xtuple_0(X0),X1) != X0) & (k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0))) | k6_relat_1(k10_xtuple_0(X0),X1) = X0)) & v1_relat_1(X1)) & v1_relat_1(X0)) [nnf transformation 14]
25. ? [X0] : (? [X1] : ((k2_relat_1(X0) != k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0))) | k6_relat_1(k10_xtuple_0(X0),X1) != X0) & (k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0))) | k6_relat_1(k10_xtuple_0(X0),X1) = X0) & v1_relat_1(X1)) & v1_relat_1(X0)) [flattening 24]
26. ? [X0] : (? [X1] : ((k2_relat_1(X0) != k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0))) | k6_relat_1(k10_xtuple_0(X0),X1) != X0) & (k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0))) | k6_relat_1(k10_xtuple_0(X0),X1) = X0) & v1_relat_1(X1)) & v1_relat_1(X0)) => (? [X1] : ((k2_relat_1(sK0) != k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 != k6_relat_1(k10_xtuple_0(sK0),X1)) & (k2_relat_1(sK0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 = k6_relat_1(k10_xtuple_0(sK0),X1)) & v1_relat_1(X1)) & v1_relat_1(sK0)) [choice axiom]
27. ? [X1] : ((k2_relat_1(sK0) != k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 != k6_relat_1(k10_xtuple_0(sK0),X1)) & (k2_relat_1(sK0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 = k6_relat_1(k10_xtuple_0(sK0),X1)) & v1_relat_1(X1)) => ((k2_relat_1(sK0) != k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 != k6_relat_1(k10_xtuple_0(sK0),sK1)) & (k2_relat_1(sK0) = k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 = k6_relat_1(k10_xtuple_0(sK0),sK1)) & v1_relat_1(sK1)) [choice axiom]
28. ((k2_relat_1(sK0) != k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 != k6_relat_1(k10_xtuple_0(sK0),sK1)) & (k2_relat_1(sK0) = k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 = k6_relat_1(k10_xtuple_0(sK0),sK1)) & v1_relat_1(sK1)) & v1_relat_1(sK0) [skolemisation 25,27,26]
29. ! [X0] : (! [X1] : (((k6_relat_1(k10_xtuple_0(X0),X1) = X0 | k2_relat_1(X0) != k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0)))) & (k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0))) | k6_relat_1(k10_xtuple_0(X0),X1) != X0)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [nnf transformation 19]
30. v1_relat_1(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
31. v1_relat_1(sK1) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
32. k2_relat_1(sK0) = k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 = k6_relat_1(k10_xtuple_0(sK0),sK1) [cnf transformation 28] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
33. k2_relat_1(sK0) != k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) | sK0 != k6_relat_1(k10_xtuple_0(sK0),sK1) [cnf transformation 28] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. k6_relat_1(k10_xtuple_0(X0),X1) != X0 | k2_relat_1(X0) = k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0))) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
39. k2_relat_1(X0) != k5_relat_1(k2_relat_1(X1),k9_xtuple_0(k2_relat_1(X0))) | k6_relat_1(k10_xtuple_0(X0),X1) = X0 | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. 1 <=> sK0 = k6_relat_1(k10_xtuple_0(sK0),sK1) [avatar definition]
46. sK0 = k6_relat_1(k10_xtuple_0(sK0),sK1) <- (1) [avatar component clause 45] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
47. sK0 != k6_relat_1(k10_xtuple_0(sK0),sK1) <- (~1) [avatar component clause 45] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
49. 2 <=> k2_relat_1(sK0) = k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) [avatar definition]
50. k2_relat_1(sK0) = k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) <- (2) [avatar component clause 49] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
51. k2_relat_1(sK0) != k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) <- (~2) [avatar component clause 49] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
52. ~1 | ~2 [avatar split clause 33,49,45]
53. 1 | 2 [avatar split clause 32,49,45]
55. 3 <=> v1_relat_1(sK1) [avatar definition]
57. v1_relat_1(sK1) <- (3) [avatar component clause 55] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
58. 3 [avatar split clause 31,55]
60. 4 <=> v1_relat_1(sK0) [avatar definition]
62. v1_relat_1(sK0) <- (4) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
63. 4 [avatar split clause 30,60]
111. sK0 != sK0 | k2_relat_1(sK0) = k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (1) [superposition 38,46] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
116. k2_relat_1(sK0) = k5_relat_1(k2_relat_1(sK1),k9_xtuple_0(k2_relat_1(sK0))) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (1) [trivial inequality removal 111] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
120. ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (1, ~2) [subsumption resolution 116,51] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
121. ~v1_relat_1(sK0) <- (1, ~2, 3) [subsumption resolution 120,57] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
122. $false <- (1, ~2, 3, 4) [subsumption resolution 121,62] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
123. ~1 | 2 | ~3 | ~4 [avatar contradiction clause 122]
136. k2_relat_1(sK0) != k2_relat_1(sK0) | sK0 = k6_relat_1(k10_xtuple_0(sK0),sK1) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2) [superposition 39,50] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
143. sK0 = k6_relat_1(k10_xtuple_0(sK0),sK1) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2) [trivial inequality removal 136] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
147. ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 143,47] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
148. ~v1_relat_1(sK0) <- (~1, 2, 3) [subsumption resolution 147,57] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
149. $false <- (~1, 2, 3, 4) [subsumption resolution 148,62] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
150. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 149]
151. $false [avatar sat refutation 52,53,58,63,123,150] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:101,thDist:-101}
% SZS output end Proof for group_9__t120_group_9
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 14
% Generated clauses: 90
% Active clauses: 23
% Passive clauses: 25
% Final active clauses: 18
% Final passive clauses: 2
% Main loop iterations started: 23
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Trivial inequalities: 6
% Fw subsumption resolutions: 14
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 1
% Unit resulting resolution: 3
% Forward superposition: 13
% Backward superposition: 13
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 2
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

