i: [1,45,0,2,1,1,0,0] t37_funct_1
a: [3,78,0,2,1,45]
s: 78
i: [1,44,0,2,1,1,0,0] t37_funct_1
a: [3,83,0,2,1,44]
s: 83
i: [1,43,0,2,1,1,0,0] t37_funct_1
a: [3,88,0,2,1,43]
s: 88
i: [1,46,0,16,2,1,0,0] t37_funct_1
a: [3,69,0,8,1,46]
s: 69
i: [1,62,0,2,1,0,0,1] rc1_funct_1
a: [3,93,0,2,1,62]
s: 93
i: [1,61,0,2,1,0,0,1] rc1_funct_1
a: [3,98,0,2,1,61]
s: 98
i: [1,65,0,2,1,0,0,1] rc2_funct_1
a: [3,103,0,2,1,65]
s: 103
i: [1,64,0,2,1,0,0,1] rc2_funct_1
a: [3,108,0,2,1,64]
s: 108
i: [1,63,0,2,1,0,0,1] rc2_funct_1
a: [3,113,0,2,1,63]
s: 113
i: [1,57,0,3,1,0,0,2] reflexivity_r1_tarski
s: 57
i: [1,56,0,4,1,0,0,255] existence_m1_subset_1
s: 56
i: [1,58,0,4,1,0,0,1] dt_k3_relat_1
s: 58
i: [1,50,0,7,3,0,0,1] dt_k2_funct_1
s: 50
i: [1,51,0,7,3,0,0,1] dt_k2_funct_1
s: 51
i: [1,54,0,7,2,0,0,255] t3_subset
s: 54
i: [1,55,0,7,2,0,0,255] t3_subset
s: 55
d: [2,115,1,5,1,0,40,54,56]
s: 115
i: [1,49,0,8,3,0,0,255] cc2_relat_1
s: 49
d: [2,117,1,6,2,0,40,49,56]
s: 117
i: [1,52,0,12,4,0,0,1] t33_funct_1
s: 52
i: [1,53,0,12,4,0,0,1] t33_funct_1
s: 53
d: [2,118,1,7,3,0,40,49,55]
s: 118
i: [1,59,0,12,5,0,0,1] fc2_funct_1
s: 59
i: [1,60,0,12,5,0,0,1] fc2_funct_1
s: 60
i: [1,47,0,16,4,0,0,1] t28_relat_1
s: 47
i: [1,48,0,16,4,0,0,1] t27_relat_1
s: 48
d: [2,121,1,17,4,1,44,69,48]
d: [2,122,1,15,3,2,21,121,88]
a: [3,126,1,3,1,122]
s: 126
d: [2,136,2,0,0,3,52,88,83,126,50]
e: 136
d: [2,137,2,4,2,1,40,126,50]
d: [2,139,2,2,1,2,21,137,83]
d: [2,140,2,0,0,3,21,139,88]
e: 140
r: 126
a: [3,125,1,3,1,122]
s: 125
a: [3,130,1,6,1,122]
s: 130
d: [2,142,2,11,4,1,44,130,53]
d: [2,143,2,6,3,1,21,142,57]
d: [2,144,2,4,2,2,21,143,78]
d: [2,145,2,2,1,3,21,144,83]
d: [2,146,2,0,0,4,21,145,88]
e: 146
r: 130
a: [3,129,1,6,1,122]
s: 129
a: [3,134,1,6,1,122]
s: 134
d: [2,159,2,0,0,4,52,88,83,78,134,52]
e: 159
d: [2,160,2,11,4,1,44,134,52]
d: [2,161,2,6,3,1,19,160]
d: [2,163,2,4,2,2,21,161,78]
d: [2,164,2,2,1,3,21,163,83]
d: [2,165,2,0,0,4,21,164,88]
e: 165
r: 69
r: 134
d: [2,148,2,7,2,1,40,129,118]
a: [3,157,2,3,1,148]
s: 157
a: [3,68,0,8,1,46]
s: 68
a: [3,133,1,6,1,122]
s: 133
a: [3,73,0,8,1,46]
s: 73
d: [2,177,1,16,4,1,44,73,47]
d: [2,178,1,11,3,1,19,177]
d: [2,179,1,8,2,2,21,178,125]
d: [2,180,1,6,1,3,21,179,88]
a: [3,184,1,6,1,180]
s: 184
d: [2,186,2,11,4,1,44,184,53]
d: [2,187,2,6,3,1,21,186,57]
d: [2,188,2,4,2,2,21,187,78]
d: [2,189,2,2,1,3,21,188,83]
d: [2,190,2,0,0,4,21,189,88]
e: 190
f: 190,180,165,146,140,122,43,44,45,46
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_1__t37_funct_1
% SZS output start Proof for funct_1__t37_funct_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k10_xtuple_0(X0) = k10_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) & k9_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) = k10_xtuple_0(X0)))) [input t37_funct_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k10_xtuple_0(X0) = k10_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) & k9_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) = k10_xtuple_0(X0)))) [negated conjecture 1]
6. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0)))) [input dt_k2_funct_1]
14. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
15. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(k10_xtuple_0(X0),k9_xtuple_0(X1)) => k9_xtuple_0(k3_relat_1(X0,X1)) = k9_xtuple_0(X0)))) [input t27_relat_1]
16. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(k9_xtuple_0(X0),k10_xtuple_0(X1)) => k10_xtuple_0(X0) = k10_xtuple_0(k3_relat_1(X1,X0))))) [input t28_relat_1]
17. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))))) [input t33_funct_1]
20. ! [X0] : r1_tarski(X0,X0) [rectify 14]
21. ? [X0] : (((k10_xtuple_0(X0) != k10_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) | k9_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) != k10_xtuple_0(X0)) & v2_funct_1(X0)) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
22. ? [X0] : ((k10_xtuple_0(X0) != k10_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) | k9_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) != k10_xtuple_0(X0)) & v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 21]
23. ! [X0] : (! [X1] : ((k10_xtuple_0(X0) = k10_xtuple_0(k3_relat_1(X1,X0)) | ~r1_tarski(k9_xtuple_0(X0),k10_xtuple_0(X1))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 16]
24. ! [X0] : (! [X1] : (k10_xtuple_0(X0) = k10_xtuple_0(k3_relat_1(X1,X0)) | ~r1_tarski(k9_xtuple_0(X0),k10_xtuple_0(X1)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 23]
25. ! [X0] : (! [X1] : ((k9_xtuple_0(k3_relat_1(X0,X1)) = k9_xtuple_0(X0) | ~r1_tarski(k10_xtuple_0(X0),k9_xtuple_0(X1))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 15]
26. ! [X0] : (! [X1] : (k9_xtuple_0(k3_relat_1(X0,X1)) = k9_xtuple_0(X0) | ~r1_tarski(k10_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 25]
28. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 6]
29. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 28]
30. ! [X0] : (((k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 17]
31. ! [X0] : ((k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 30]
34. ? [X0] : ((k10_xtuple_0(X0) != k10_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) | k9_xtuple_0(k3_relat_1(k2_funct_1(X0),X0)) != k10_xtuple_0(X0)) & v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ((k10_xtuple_0(sK0) != k10_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0)) | k10_xtuple_0(sK0) != k9_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0))) & v2_funct_1(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
35. (k10_xtuple_0(sK0) != k10_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0)) | k10_xtuple_0(sK0) != k9_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0))) & v2_funct_1(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 22,34]
43. v1_relat_1(sK0) [cnf transformation 35] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
44. v1_funct_1(sK0) [cnf transformation 35] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. v2_funct_1(sK0) [cnf transformation 35] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. k10_xtuple_0(sK0) != k10_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0)) | k10_xtuple_0(sK0) != k9_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0)) [cnf transformation 35] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. k10_xtuple_0(X0) = k10_xtuple_0(k3_relat_1(X1,X0)) | ~r1_tarski(k9_xtuple_0(X0),k10_xtuple_0(X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
48. k9_xtuple_0(k3_relat_1(X0,X1)) = k9_xtuple_0(X0) | ~r1_tarski(k10_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. v1_relat_1(k2_funct_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. r1_tarski(X0,X0) [cnf transformation 20] {a:0,w:3,wCS:1200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
67. 1 <=> k10_xtuple_0(sK0) = k9_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0)) [avatar definition]
69. k10_xtuple_0(sK0) != k9_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0)) <- (~1) [avatar component clause 67] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. 2 <=> k10_xtuple_0(sK0) = k10_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0)) [avatar definition]
73. k10_xtuple_0(sK0) != k10_xtuple_0(k3_relat_1(k2_funct_1(sK0),sK0)) <- (~2) [avatar component clause 71] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
74. ~1 | ~2 [avatar split clause 46,71,67]
76. 3 <=> v2_funct_1(sK0) [avatar definition]
78. v2_funct_1(sK0) <- (3) [avatar component clause 76] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. 3 [avatar split clause 45,76]
81. 4 <=> v1_funct_1(sK0) [avatar definition]
83. v1_funct_1(sK0) <- (4) [avatar component clause 81] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. 4 [avatar split clause 44,81]
86. 5 <=> v1_relat_1(sK0) [avatar definition]
88. v1_relat_1(sK0) <- (5) [avatar component clause 86] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
89. 5 [avatar split clause 43,86]
121. k10_xtuple_0(sK0) != k9_xtuple_0(k2_funct_1(sK0)) | ~r1_tarski(k10_xtuple_0(k2_funct_1(sK0)),k9_xtuple_0(sK0)) | ~v1_relat_1(sK0) | ~v1_relat_1(k2_funct_1(sK0)) <- (~1) [superposition 69,48] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
122. k10_xtuple_0(sK0) != k9_xtuple_0(k2_funct_1(sK0)) | ~r1_tarski(k10_xtuple_0(k2_funct_1(sK0)),k9_xtuple_0(sK0)) | ~v1_relat_1(k2_funct_1(sK0)) <- (~1, 5) [subsumption resolution 121,88] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
124. 11 <=> v1_relat_1(k2_funct_1(sK0)) [avatar definition]
125. v1_relat_1(k2_funct_1(sK0)) <- (11) [avatar component clause 124] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. ~v1_relat_1(k2_funct_1(sK0)) <- (~11) [avatar component clause 124] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. 12 <=> r1_tarski(k10_xtuple_0(k2_funct_1(sK0)),k9_xtuple_0(sK0)) [avatar definition]
130. ~r1_tarski(k10_xtuple_0(k2_funct_1(sK0)),k9_xtuple_0(sK0)) <- (~12) [avatar component clause 128] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
132. 13 <=> k10_xtuple_0(sK0) = k9_xtuple_0(k2_funct_1(sK0)) [avatar definition]
134. k10_xtuple_0(sK0) != k9_xtuple_0(k2_funct_1(sK0)) <- (~13) [avatar component clause 132] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
135. ~11 | ~12 | ~13 | 1 | ~5 [avatar split clause 122,86,67,132,128,124]
137. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~11) [resolution 126,50] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
139. ~v1_relat_1(sK0) <- (4, ~11) [subsumption resolution 137,83] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
140. $false <- (4, 5, ~11) [subsumption resolution 139,88] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
141. ~4 | ~5 | 11 [avatar contradiction clause 140]
142. ~r1_tarski(k9_xtuple_0(sK0),k9_xtuple_0(sK0)) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~12) [superposition 130,53] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
143. ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~12) [subsumption resolution 142,57] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
144. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3, ~12) [subsumption resolution 143,78] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
145. ~v1_relat_1(sK0) <- (3, 4, ~12) [subsumption resolution 144,83] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
146. $false <- (3, 4, 5, ~12) [subsumption resolution 145,88] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
147. ~3 | ~4 | ~5 | 12 [avatar contradiction clause 146]
160. k10_xtuple_0(sK0) != k10_xtuple_0(sK0) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~13) [superposition 134,52] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
161. ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~13) [trivial inequality removal 160] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
163. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3, ~13) [subsumption resolution 161,78] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
164. ~v1_relat_1(sK0) <- (3, 4, ~13) [subsumption resolution 163,83] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
165. $false <- (3, 4, 5, ~13) [subsumption resolution 164,88] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
166. ~3 | ~4 | ~5 | 13 [avatar contradiction clause 165]
177. k10_xtuple_0(sK0) != k10_xtuple_0(sK0) | ~r1_tarski(k9_xtuple_0(sK0),k10_xtuple_0(k2_funct_1(sK0))) | ~v1_relat_1(k2_funct_1(sK0)) | ~v1_relat_1(sK0) <- (~2) [superposition 73,47] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
178. ~r1_tarski(k9_xtuple_0(sK0),k10_xtuple_0(k2_funct_1(sK0))) | ~v1_relat_1(k2_funct_1(sK0)) | ~v1_relat_1(sK0) <- (~2) [trivial inequality removal 177] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
179. ~r1_tarski(k9_xtuple_0(sK0),k10_xtuple_0(k2_funct_1(sK0))) | ~v1_relat_1(sK0) <- (~2, 11) [subsumption resolution 178,125] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
180. ~r1_tarski(k9_xtuple_0(sK0),k10_xtuple_0(k2_funct_1(sK0))) <- (~2, 5, 11) [subsumption resolution 179,88] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
182. 16 <=> r1_tarski(k9_xtuple_0(sK0),k10_xtuple_0(k2_funct_1(sK0))) [avatar definition]
184. ~r1_tarski(k9_xtuple_0(sK0),k10_xtuple_0(k2_funct_1(sK0))) <- (~16) [avatar component clause 182] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
185. ~16 | 2 | ~5 | ~11 [avatar split clause 180,124,86,71,182]
186. ~r1_tarski(k9_xtuple_0(sK0),k9_xtuple_0(sK0)) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~16) [superposition 184,53] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
187. ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~16) [subsumption resolution 186,57] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
188. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3, ~16) [subsumption resolution 187,78] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
189. ~v1_relat_1(sK0) <- (3, 4, ~16) [subsumption resolution 188,83] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
190. $false <- (3, 4, 5, ~16) [subsumption resolution 189,88] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
191. ~3 | ~4 | ~5 | 16 [avatar contradiction clause 190]
192. $false [avatar sat refutation 74,79,84,89,135,141,147,166,185,191] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:101,thDist:-101}
% SZS output end Proof for funct_1__t37_funct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 84
% Active clauses: 35
% Passive clauses: 36
% Final active clauses: 31
% Main loop iterations started: 36
% 
% >>> Simplifying Inferences
% Trivial inequalities: 2
% Fw subsumption resolutions: 22
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 8
% Unit resulting resolution: 2
% Forward superposition: 6
% Backward superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 22
% SAT solver unit clauses: 8
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

