i: [1,80,0,3,1,1,0,0] t6_arytm_1
a: [3,113,0,3,1,80]
s: 113
i: [1,79,0,3,1,1,0,0] t6_arytm_1
a: [3,118,0,3,1,79]
s: 118
r: 113
i: [1,78,0,3,1,1,0,0] t6_arytm_1
a: [3,123,0,3,1,78]
s: 123
i: [1,77,0,3,1,1,0,0] t6_arytm_1
a: [3,128,0,3,1,77]
s: 128
d: [2,158,0,3,1,2,23,113,118]
a: [3,162,0,3,1,158]
s: 162
i: [1,96,0,3,1,0,0,255] reflexivity_r1_tarski
s: 96
i: [1,81,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,133,0,3,1,81]
s: 133
i: [1,82,0,3,1,0,0,1] dt_k11_arytm_3
a: [3,138,0,3,1,82]
s: 138
i: [1,107,0,11,3,0,0,255] d8_arytm_2
a: [3,147,0,3,1,107]
s: 147
i: [1,95,0,4,1,0,0,1] existence_m1_subset_1
s: 95
i: [1,83,0,5,1,0,0,2] dt_k1_arytm_2
a: [3,143,0,5,1,83]
s: 143
i: [1,90,0,6,2,0,0,255] t1_subset
s: 90
i: [1,97,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 97
i: [1,84,0,7,2,0,0,255] dt_k4_arytm_2
s: 84
i: [1,92,0,7,2,0,0,255] t3_subset
s: 92
i: [1,91,0,8,3,0,0,255] t2_subset
s: 91
d: [2,167,1,5,2,1,40,91,138]
a: [3,180,1,2,1,167]
s: 180
a: [3,151,0,8,2,107]
s: 151
d: [2,168,1,5,2,1,40,91,147]
a: [3,189,1,2,1,168]
s: 189
i: [1,108,0,11,3,0,0,255] d8_arytm_2
a: [3,155,0,8,2,108]
s: 155
i: [1,85,0,9,2,0,0,255] dt_k3_arytm_2
s: 85
d: [2,169,1,9,2,1,40,91,143]
a: [3,194,1,5,1,169]
s: 194
i: [1,94,0,9,3,0,0,255] t5_subset
s: 94
d: [2,217,1,6,2,1,40,94,143]
a: [3,225,1,3,1,217]
s: 225
i: [1,93,0,10,3,0,0,255] t4_subset
s: 93
i: [1,104,0,11,3,0,0,2] dt_k7_arytm_2
s: 104
d: [2,209,2,5,1,1,40,194,97]
a: [3,214,2,5,1,209]
s: 214
i: [1,106,0,12,4,0,0,1] connectedness_r1_arytm_2
s: 106
d: [2,238,1,9,3,1,40,106,123]
s: 238
i: [1,102,0,13,3,0,0,255] commutativity_k5_arytm_2
s: 102
i: [1,103,0,13,3,0,0,255] dt_k5_arytm_2
s: 103
d: [2,240,2,6,2,2,40,238,147]
d: [2,246,2,3,1,3,21,240,162]
a: [3,250,2,3,1,246]
s: 250
i: [1,105,0,13,3,0,0,2] commutativity_k7_arytm_2
s: 105
d: [2,243,2,6,2,1,40,238,123]
d: [2,245,2,3,1,1,18,243]
a: [3,255,2,3,1,245]
s: 255
i: [1,101,0,14,4,0,0,255] existence_m2_subset_1
s: 101
i: [1,98,0,15,5,0,0,255] dt_m2_subset_1
s: 98
r: 225
d: [2,244,2,8,2,1,40,238,95]
a: [3,260,2,4,1,244]
s: 260
i: [1,99,0,15,5,0,0,255] redefinition_m2_subset_1
s: 99
d: [2,242,2,11,3,1,40,238,84]
s: 242
i: [1,100,0,15,5,0,0,255] redefinition_m2_subset_1
s: 100
i: [1,109,0,16,4,0,0,1] t19_arytm_2
s: 109
a: [3,221,1,3,1,217]
s: 221
r: 214
d: [2,333,0,11,3,0,21,109,104]
s: 333
a: [3,224,1,3,1,217]
s: 224
i: [1,88,0,22,5,0,0,255] d8_arytm_2
s: 88
d: [2,172,1,6,2,0,40,91,95]
s: 172
r: 221
r: 224
d: [2,290,1,18,5,0,40,98,85]
d: [2,294,1,13,4,1,21,290,143]
d: [2,295,1,10,3,2,21,294,225]
a: [3,299,1,2,1,295]
s: 299
s: 225
d: [2,241,2,16,4,1,40,238,104]
s: 241
d: [2,334,1,12,4,1,44,333,151]
d: [2,341,1,9,3,1,18,334]
d: [2,342,1,6,2,2,21,341,147]
s: 342
d: [2,335,1,12,4,1,44,333,155]
d: [2,340,1,9,3,1,18,335]
d: [2,343,1,6,2,2,21,340,147]
s: 343
d: [2,375,2,0,0,4,52,162,123,343]
e: 375
d: [2,379,2,3,1,3,40,343,123]
d: [2,382,2,0,0,4,21,379,162]
e: 382
d: [2,164,0,3,1,2,22,128,118]
f: 382,164,158,108,77,78,79,80
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t6_arytm_1
% SZS output start Proof for arytm_1__t6_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => r1_arytm_2(X0,X1)))) [input t6_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => r1_arytm_2(X0,X1)))) [negated conjecture 1]
7. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (~(k7_arytm_2(X0,X1) != k4_arytm_2(k5_arytm_2(k3_arytm_2(X0),k3_arytm_2(X1))) & k11_arytm_3 != X0 & k11_arytm_3 != X1) & (k11_arytm_3 = X0 => k7_arytm_2(X0,X1) = X1) & (k11_arytm_3 = X1 => k7_arytm_2(X0,X1) = X0)))) [input d8_arytm_2]
17. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2)) [input dt_k7_arytm_2]
25. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (k7_arytm_2(X1,X2) = X0 => r1_arytm_2(X2,X0))))) [input t19_arytm_2]
38. ? [X0] : (? [X1] : ((~r1_arytm_2(X0,X1) & k11_arytm_3 = X0) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
39. ? [X0] : (? [X1] : (~r1_arytm_2(X0,X1) & k11_arytm_3 = X0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 38]
42. ! [X0] : (! [X1] : (((k7_arytm_2(X0,X1) = k4_arytm_2(k5_arytm_2(k3_arytm_2(X0),k3_arytm_2(X1))) | k11_arytm_3 = X0 | k11_arytm_3 = X1) & (k7_arytm_2(X0,X1) = X1 | k11_arytm_3 != X0) & (k7_arytm_2(X0,X1) = X0 | k11_arytm_3 != X1)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 7]
43. ! [X0] : (! [X1] : (! [X2] : ((r1_arytm_2(X2,X0) | k7_arytm_2(X1,X2) != X0) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 25]
44. ! [X0] : (! [X1] : (! [X2] : (r1_arytm_2(X2,X0) | k7_arytm_2(X1,X2) != X0 | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 43]
63. ! [X0,X1] : (m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 17]
64. ! [X0,X1] : (m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 63]
69. ? [X0] : (? [X1] : (~r1_arytm_2(X0,X1) & k11_arytm_3 = X0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (~r1_arytm_2(sK0,X1) & k11_arytm_3 = sK0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
70. ? [X1] : (~r1_arytm_2(sK0,X1) & k11_arytm_3 = sK0 & m1_subset_1(X1,k2_arytm_2)) => (~r1_arytm_2(sK0,sK1) & k11_arytm_3 = sK0 & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
71. (~r1_arytm_2(sK0,sK1) & k11_arytm_3 = sK0 & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 39,70,69]
77. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 71] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
78. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 71] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
79. k11_arytm_3 = sK0 [cnf transformation 71] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
80. ~r1_arytm_2(sK0,sK1) [cnf transformation 71] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
86. k7_arytm_2(X0,X1) = X0 | k11_arytm_3 != X1 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 42] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
89. r1_arytm_2(X2,X0) | k7_arytm_2(X1,X2) != X0 | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 44] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
104. m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 64] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
108. k7_arytm_2(X0,k11_arytm_3) = X0 | ~m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [equality resolution 86] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
109. r1_arytm_2(X2,k7_arytm_2(X1,X2)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k7_arytm_2(X1,X2),k2_arytm_2) [equality resolution 89] {a:0,w:16,wCS:6400,sine:1,thAx:0,allAx:1,thDist:-1}
111. 1 <=> r1_arytm_2(sK0,sK1) [avatar definition]
113. ~r1_arytm_2(sK0,sK1) <- (~1) [avatar component clause 111] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. ~1 [avatar split clause 80,111]
116. 2 <=> k11_arytm_3 = sK0 [avatar definition]
118. k11_arytm_3 = sK0 <- (2) [avatar component clause 116] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. 2 [avatar split clause 79,116]
121. 3 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
123. m1_subset_1(sK1,k2_arytm_2) <- (3) [avatar component clause 121] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. 3 [avatar split clause 78,121]
126. 4 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
128. m1_subset_1(sK0,k2_arytm_2) <- (4) [avatar component clause 126] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. 4 [avatar split clause 77,126]
146. 8 <=> m1_subset_1(k11_arytm_3,k2_arytm_2) [avatar definition]
147. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (8) [avatar component clause 146] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
154. 10 <=> ! [X0] : (k7_arytm_2(X0,k11_arytm_3) = X0 | ~m1_subset_1(X0,k2_arytm_2)) [avatar definition]
155. k7_arytm_2(X0,k11_arytm_3) = X0 | ~m1_subset_1(X0,k2_arytm_2) <- (10) [avatar component clause 154] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
156. ~8 | 10 [avatar split clause 108,154,146]
158. ~r1_arytm_2(k11_arytm_3,sK1) <- (~1, 2) [backward demodulation 113,118] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
160. 11 <=> r1_arytm_2(k11_arytm_3,sK1) [avatar definition]
162. ~r1_arytm_2(k11_arytm_3,sK1) <- (~11) [avatar component clause 160] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
163. ~11 | 1 | ~2 [avatar split clause 158,116,111,160]
164. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (2, 4) [forward demodulation 128,118] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
165. 8 | ~2 | ~4 [avatar split clause 164,126,116,146]
333. r1_arytm_2(X2,k7_arytm_2(X1,X2)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) [subsumption resolution 109,104] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
335. r1_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) <- (10) [superposition 333,155] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
340. r1_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) <- (10) [duplicate literal removal 335] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
343. ~m1_subset_1(X1,k2_arytm_2) | r1_arytm_2(k11_arytm_3,X1) <- (8, 10) [subsumption resolution 340,147] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
379. r1_arytm_2(k11_arytm_3,sK1) <- (3, 8, 10) [resolution 343,123] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
382. $false <- (3, 8, 10, ~11) [subsumption resolution 379,162] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
383. ~3 | ~8 | ~10 | 11 [avatar contradiction clause 382]
384. $false [avatar sat refutation 114,119,124,129,156,163,165,383] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:179,thDist:-179}
% SZS output end Proof for arytm_1__t6_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% >>> Preprocessing
% Introduced skolems: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 30
% Generated clauses: 171
% Active clauses: 49
% Passive clauses: 76
% Final active clauses: 44
% Final passive clauses: 25
% Main loop iterations started: 51
% 
% >>> Simplifying Inferences
% Duplicate literals: 37
% Fw subsumption resolutions: 11
% Fw demodulations: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 21
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 48
% Unit resulting resolution: 1
% Forward superposition: 18
% Backward superposition: 8
% 
% >>> AVATAR
% Split clauses: 16
% Split components: 37
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 36
% SAT solver unit clauses: 7
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

