i: [1,98,0,2,1,1,0,0] t10_isomichi
a: [3,156,0,2,1,98]
s: 156
i: [1,97,0,2,1,1,0,0] t10_isomichi
a: [3,161,0,2,1,97]
s: 161
i: [1,102,0,17,3,1,0,0] t10_isomichi
a: [3,136,0,3,1,102]
s: 136
i: [1,99,0,5,1,1,0,0] t10_isomichi
a: [3,151,0,5,1,99]
s: 151
a: [3,139,0,7,1,102]
s: 139
a: [3,143,0,7,1,102]
s: 143
i: [1,132,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,166,0,2,1,132]
s: 166
i: [1,126,0,3,1,0,0,1] reflexivity_r1_tarski
s: 126
i: [1,125,0,4,1,0,0,1] existence_m1_subset_1
s: 125
i: [1,123,0,7,2,0,0,1] t3_subset
s: 123
i: [1,124,0,7,2,0,0,1] t3_subset
s: 124
d: [2,168,1,4,1,1,40,123,151]
a: [3,173,1,4,1,168]
s: 173
i: [1,114,0,8,3,0,0,1] rc2_isomichi
s: 114
d: [2,169,1,5,1,0,40,123,125]
s: 169
i: [1,115,0,8,3,0,0,1] rc2_isomichi
s: 115
i: [1,116,0,8,3,0,0,1] rc2_isomichi
s: 116
i: [1,118,0,8,3,0,0,255] rc6_pre_topc
s: 118
i: [1,122,0,8,3,0,0,255] rc2_tops_1
s: 122
i: [1,113,0,10,3,0,0,1] rc2_isomichi
s: 113
i: [1,117,0,10,3,0,0,255] rc6_pre_topc
s: 117
i: [1,119,0,10,3,0,0,255] rc1_tops_1
s: 119
d: [2,176,1,9,3,0,40,113,123]
s: 176
i: [1,120,0,10,3,0,0,255] rc10_pre_topc
s: 120
d: [2,177,1,9,3,0,40,117,123]
s: 177
i: [1,121,0,10,3,0,0,255] rc2_tops_1
s: 121
i: [1,127,0,14,4,0,0,255] fc1_tops_1
s: 127
d: [2,182,1,9,3,1,40,127,151]
d: [2,194,1,7,2,2,21,182,156]
d: [2,195,1,5,1,3,21,194,161]
a: [3,199,1,5,1,195]
s: 199
i: [1,128,0,14,3,0,0,1] dt_k2_pre_topc
s: 128
d: [2,178,1,9,3,0,40,119,123]
s: 178
i: [1,129,0,14,3,0,0,1] dt_k1_tops_1
s: 129
i: [1,103,0,15,5,0,0,255] cc5_isomichi
s: 103
d: [2,210,1,10,4,1,40,103,151]
d: [2,224,1,8,3,2,21,210,156]
d: [2,225,1,6,2,3,21,224,161]
a: [3,233,1,3,1,225]
s: 233
i: [1,104,0,15,5,0,0,1] cc2_isomichi
s: 104
d: [2,179,1,9,3,0,40,120,123]
s: 179
i: [1,105,0,15,5,0,0,1] cc2_isomichi
s: 105
i: [1,106,0,15,5,0,0,255] cc1_isomichi
s: 106
d: [2,272,1,10,4,1,40,106,151]
d: [2,287,1,8,3,2,21,272,156]
d: [2,288,1,6,2,3,21,287,161]
a: [3,296,1,3,1,288]
s: 296
i: [1,107,0,15,5,0,0,255] cc4_isomichi
s: 107
a: [3,228,1,3,1,225]
s: 228
i: [1,130,0,16,3,0,0,255] projectivity_k1_tops_1
s: 130
i: [1,131,0,16,3,0,0,255] projectivity_k2_pre_topc
s: 131
d: [2,274,1,16,6,0,40,106,117]
d: [2,282,1,12,4,0,18,274]
d: [2,298,1,8,3,0,21,282,118]
s: 298
i: [1,108,0,18,6,0,0,1] cc3_isomichi
s: 108
d: [2,277,1,16,6,0,40,106,121]
d: [2,279,1,12,4,0,18,277]
d: [2,299,1,8,3,0,21,279,122]
s: 299
i: [1,109,0,19,5,0,0,1] t8_isomichi
s: 109
i: [1,110,0,19,5,0,0,1] t8_isomichi
s: 110
a: [3,292,1,3,1,288]
s: 292
i: [1,111,0,19,5,0,0,1] t9_isomichi
s: 111
d: [2,304,1,16,6,0,40,107,113]
d: [2,314,1,12,4,0,18,304]
d: [2,321,1,8,3,0,21,314,116]
s: 321
i: [1,112,0,19,5,0,0,1] t9_isomichi
s: 112
d: [2,364,2,0,0,5,52,161,156,228,151,143,112]
e: 364
d: [2,366,1,12,4,1,40,112,143]
d: [2,371,1,9,3,2,21,366,228]
d: [2,372,1,4,2,3,21,371,151]
d: [2,373,1,2,1,4,21,372,156]
d: [2,374,1,0,0,5,21,373,161]
e: 374
r: 228
r: 233
r: 136
r: 139
r: 292
a: [3,135,0,3,1,102]
s: 135
a: [3,229,1,3,1,225]
s: 229
a: [3,140,0,7,1,102]
s: 140
r: 140
r: 143
a: [3,232,1,3,1,225]
s: 232
s: 139
a: [3,144,0,7,1,102]
s: 144
d: [2,408,2,0,0,5,52,161,156,229,151,144,111]
e: 408
d: [2,409,1,12,4,1,40,144,111]
d: [2,411,1,9,3,2,21,409,229]
d: [2,412,1,4,2,3,21,411,151]
d: [2,413,1,2,1,4,21,412,156]
d: [2,414,1,0,0,5,21,413,161]
e: 414
i: [1,101,0,10,2,1,0,0] t10_isomichi
i: [1,100,0,10,2,1,0,0] t10_isomichi
d: [2,377,1,7,3,2,21,366,151]
d: [2,378,1,5,2,3,21,377,156]
d: [2,379,1,3,1,4,21,378,161]
d: [2,255,1,10,4,1,40,105,151]
d: [2,389,1,8,3,2,21,255,156]
d: [2,390,1,6,2,3,21,389,161]
d: [2,238,1,10,4,1,40,104,151]
d: [2,392,1,8,3,2,21,238,156]
d: [2,393,1,6,2,3,21,392,161]
d: [2,333,1,13,5,1,40,108,151]
d: [2,395,1,11,4,2,21,333,156]
d: [2,396,1,9,3,3,21,395,161]
d: [2,398,1,12,4,1,40,140,109]
d: [2,399,1,7,3,2,21,398,151]
d: [2,400,1,5,2,3,21,399,156]
d: [2,401,1,3,1,4,21,400,161]
d: [2,403,1,12,4,1,40,139,110]
d: [2,404,1,7,3,2,21,403,151]
d: [2,405,1,5,2,3,21,404,156]
d: [2,406,1,3,1,4,21,405,161]
f: 414,406,401,396,393,390,379,97,98,99,100,101,102
% Refutation found. Thanks to Tanya!
% SZS status Theorem for isomichi__t10_isomichi
% SZS output start Proof for isomichi__t10_isomichi
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v4_tops_1(X1,X0) <=> (r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) & r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1))))) [input t10_isomichi]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v4_tops_1(X1,X0) <=> (r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) & r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1))))) [negated conjecture 1]
4. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v4_tops_1(X1,X0) => (v2_isomichi(X1,X0) & v1_isomichi(X1,X0))))) [input cc2_isomichi]
5. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ((v2_isomichi(X1,X0) & v1_isomichi(X1,X0)) => v4_tops_1(X1,X0)))) [input cc3_isomichi]
29. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v1_isomichi(X1,X0) <=> r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1)))) [input t8_isomichi]
30. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v2_isomichi(X1,X0) <=> r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1)))))) [input t9_isomichi]
39. ? [X0] : (? [X1] : ((v4_tops_1(X1,X0) <~> (r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) & r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1))) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & (l1_pre_topc(X0) & v2_pre_topc(X0))) [ennf transformation 2]
40. ? [X0] : (? [X1] : ((v4_tops_1(X1,X0) <~> (r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) & r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1))) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 39]
43. ! [X0] : (! [X1] : (((v2_isomichi(X1,X0) & v1_isomichi(X1,X0)) | ~v4_tops_1(X1,X0)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 4]
44. ! [X0] : (! [X1] : ((v2_isomichi(X1,X0) & v1_isomichi(X1,X0)) | ~v4_tops_1(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 43]
49. ! [X0] : (! [X1] : ((v4_tops_1(X1,X0) | (~v2_isomichi(X1,X0) | ~v1_isomichi(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 5]
50. ! [X0] : (! [X1] : (v4_tops_1(X1,X0) | ~v2_isomichi(X1,X0) | ~v1_isomichi(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 49]
51. ! [X0] : (! [X1] : ((v1_isomichi(X1,X0) <=> r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 29]
52. ! [X0] : (! [X1] : ((v1_isomichi(X1,X0) <=> r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 51]
53. ! [X0] : (! [X1] : ((v2_isomichi(X1,X0) <=> r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1)))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 30]
54. ! [X0] : (! [X1] : ((v2_isomichi(X1,X0) <=> r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1)))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 53]
75. ? [X0] : (? [X1] : ((((~r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) | ~r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1)) | ~v4_tops_1(X1,X0)) & ((r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) & r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1)) | v4_tops_1(X1,X0))) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) [nnf transformation 40]
76. ? [X0] : (? [X1] : ((~r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) | ~r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1) | ~v4_tops_1(X1,X0)) & ((r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) & r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1)) | v4_tops_1(X1,X0)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 75]
77. ? [X0] : (? [X1] : ((~r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) | ~r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1) | ~v4_tops_1(X1,X0)) & ((r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) & r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1)) | v4_tops_1(X1,X0)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) => (? [X1] : ((~r1_tarski(X1,k2_pre_topc(sK0,k1_tops_1(sK0,X1))) | ~r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,X1)),X1) | ~v4_tops_1(X1,sK0)) & ((r1_tarski(X1,k2_pre_topc(sK0,k1_tops_1(sK0,X1))) & r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,X1)),X1)) | v4_tops_1(X1,sK0)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0)) [choice axiom]
78. ? [X1] : ((~r1_tarski(X1,k2_pre_topc(sK0,k1_tops_1(sK0,X1))) | ~r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,X1)),X1) | ~v4_tops_1(X1,sK0)) & ((r1_tarski(X1,k2_pre_topc(sK0,k1_tops_1(sK0,X1))) & r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,X1)),X1)) | v4_tops_1(X1,sK0)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => ((~r1_tarski(sK1,k2_pre_topc(sK0,k1_tops_1(sK0,sK1))) | ~r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,sK1)),sK1) | ~v4_tops_1(sK1,sK0)) & ((r1_tarski(sK1,k2_pre_topc(sK0,k1_tops_1(sK0,sK1))) & r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,sK1)),sK1)) | v4_tops_1(sK1,sK0)) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
79. ((~r1_tarski(sK1,k2_pre_topc(sK0,k1_tops_1(sK0,sK1))) | ~r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,sK1)),sK1) | ~v4_tops_1(sK1,sK0)) & ((r1_tarski(sK1,k2_pre_topc(sK0,k1_tops_1(sK0,sK1))) & r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,sK1)),sK1)) | v4_tops_1(sK1,sK0)) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) [skolemisation 76,78,77]
80. ! [X0] : (! [X1] : (((v1_isomichi(X1,X0) | ~r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1)) & (r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1) | ~v1_isomichi(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [nnf transformation 52]
81. ! [X0] : (! [X1] : (((v2_isomichi(X1,X0) | ~r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1)))) & (r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) | ~v2_isomichi(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [nnf transformation 54]
97. v2_pre_topc(sK0) [cnf transformation 79] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
98. l1_pre_topc(sK0) [cnf transformation 79] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
99. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 79] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,sK1)),sK1) | v4_tops_1(sK1,sK0) [cnf transformation 79] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
101. r1_tarski(sK1,k2_pre_topc(sK0,k1_tops_1(sK0,sK1))) | v4_tops_1(sK1,sK0) [cnf transformation 79] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
102. ~r1_tarski(sK1,k2_pre_topc(sK0,k1_tops_1(sK0,sK1))) | ~r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,sK1)),sK1) | ~v4_tops_1(sK1,sK0) [cnf transformation 79] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
104. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v4_tops_1(X1,X0) | v1_isomichi(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 44] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
105. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v4_tops_1(X1,X0) | v2_isomichi(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 44] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
108. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v2_isomichi(X1,X0) | ~v1_isomichi(X1,X0) | v4_tops_1(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 50] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
109. r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1) | ~v1_isomichi(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 80] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
110. ~r1_tarski(k1_tops_1(X0,k2_pre_topc(X0,X1)),X1) | v1_isomichi(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 80] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
111. r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) | ~v2_isomichi(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 81] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
112. ~r1_tarski(X1,k2_pre_topc(X0,k1_tops_1(X0,X1))) | v2_isomichi(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 81] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
134. 1 <=> v4_tops_1(sK1,sK0) [avatar definition]
138. 2 <=> r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,sK1)),sK1) [avatar definition]
139. r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,sK1)),sK1) <- (2) [avatar component clause 138] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
140. ~r1_tarski(k1_tops_1(sK0,k2_pre_topc(sK0,sK1)),sK1) <- (~2) [avatar component clause 138] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. 3 <=> r1_tarski(sK1,k2_pre_topc(sK0,k1_tops_1(sK0,sK1))) [avatar definition]
143. r1_tarski(sK1,k2_pre_topc(sK0,k1_tops_1(sK0,sK1))) <- (3) [avatar component clause 142] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. ~r1_tarski(sK1,k2_pre_topc(sK0,k1_tops_1(sK0,sK1))) <- (~3) [avatar component clause 142] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
145. ~1 | ~2 | ~3 [avatar split clause 102,142,138,134]
146. 1 | 3 [avatar split clause 101,142,134]
147. 1 | 2 [avatar split clause 100,138,134]
149. 4 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
151. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (4) [avatar component clause 149] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
152. 4 [avatar split clause 99,149]
154. 5 <=> l1_pre_topc(sK0) [avatar definition]
156. l1_pre_topc(sK0) <- (5) [avatar component clause 154] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
157. 5 [avatar split clause 98,154]
159. 6 <=> v2_pre_topc(sK0) [avatar definition]
161. v2_pre_topc(sK0) <- (6) [avatar component clause 159] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
162. 6 [avatar split clause 97,159]
227. 10 <=> v2_isomichi(sK1,sK0) [avatar definition]
229. v2_isomichi(sK1,sK0) <- (10) [avatar component clause 227] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
238. ~v4_tops_1(sK1,sK0) | v1_isomichi(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (4) [resolution 104,151] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
255. ~v4_tops_1(sK1,sK0) | v2_isomichi(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (4) [resolution 105,151] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
290. 12 <=> v1_isomichi(sK1,sK0) [avatar definition]
333. ~v2_isomichi(sK1,sK0) | ~v1_isomichi(sK1,sK0) | v4_tops_1(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (4) [resolution 108,151] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
366. v2_isomichi(sK1,sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (3) [resolution 112,143] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
377. v2_isomichi(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (3, 4) [subsumption resolution 366,151] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
378. v2_isomichi(sK1,sK0) | ~v2_pre_topc(sK0) <- (3, 4, 5) [subsumption resolution 377,156] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
379. v2_isomichi(sK1,sK0) <- (3, 4, 5, 6) [subsumption resolution 378,161] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
380. 10 | ~3 | ~4 | ~5 | ~6 [avatar split clause 379,159,154,149,142,227]
389. ~v4_tops_1(sK1,sK0) | v2_isomichi(sK1,sK0) | ~v2_pre_topc(sK0) <- (4, 5) [subsumption resolution 255,156] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
390. ~v4_tops_1(sK1,sK0) | v2_isomichi(sK1,sK0) <- (4, 5, 6) [subsumption resolution 389,161] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
391. 10 | ~1 | ~4 | ~5 | ~6 [avatar split clause 390,159,154,149,134,227]
392. ~v4_tops_1(sK1,sK0) | v1_isomichi(sK1,sK0) | ~v2_pre_topc(sK0) <- (4, 5) [subsumption resolution 238,156] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
393. ~v4_tops_1(sK1,sK0) | v1_isomichi(sK1,sK0) <- (4, 5, 6) [subsumption resolution 392,161] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
394. 12 | ~1 | ~4 | ~5 | ~6 [avatar split clause 393,159,154,149,134,290]
395. ~v2_isomichi(sK1,sK0) | ~v1_isomichi(sK1,sK0) | v4_tops_1(sK1,sK0) | ~v2_pre_topc(sK0) <- (4, 5) [subsumption resolution 333,156] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
396. ~v2_isomichi(sK1,sK0) | ~v1_isomichi(sK1,sK0) | v4_tops_1(sK1,sK0) <- (4, 5, 6) [subsumption resolution 395,161] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
397. 1 | ~12 | ~10 | ~4 | ~5 | ~6 [avatar split clause 396,159,154,149,227,290,134]
398. ~v1_isomichi(sK1,sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~2) [resolution 140,109] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
399. ~v1_isomichi(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~2, 4) [subsumption resolution 398,151] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
400. ~v1_isomichi(sK1,sK0) | ~v2_pre_topc(sK0) <- (~2, 4, 5) [subsumption resolution 399,156] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
401. ~v1_isomichi(sK1,sK0) <- (~2, 4, 5, 6) [subsumption resolution 400,161] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
402. ~12 | 2 | ~4 | ~5 | ~6 [avatar split clause 401,159,154,149,138,290]
403. v1_isomichi(sK1,sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (2) [resolution 139,110] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
404. v1_isomichi(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (2, 4) [subsumption resolution 403,151] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
405. v1_isomichi(sK1,sK0) | ~v2_pre_topc(sK0) <- (2, 4, 5) [subsumption resolution 404,156] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
406. v1_isomichi(sK1,sK0) <- (2, 4, 5, 6) [subsumption resolution 405,161] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
407. 12 | ~2 | ~4 | ~5 | ~6 [avatar split clause 406,159,154,149,138,290]
409. ~v2_isomichi(sK1,sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~3) [resolution 144,111] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
411. ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~3, 10) [subsumption resolution 409,229] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
412. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~3, 4, 10) [subsumption resolution 411,151] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
413. ~v2_pre_topc(sK0) <- (~3, 4, 5, 10) [subsumption resolution 412,156] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
414. $false <- (~3, 4, 5, 6, 10) [subsumption resolution 413,161] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
415. 3 | ~4 | ~5 | ~6 | ~10 [avatar contradiction clause 414]
416. $false [avatar sat refutation 145,146,147,152,157,162,380,391,394,397,402,407,415] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:197,thDist:-197}
% SZS output end Proof for isomichi__t10_isomichi
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 269
% Active clauses: 56
% Passive clauses: 119
% Final active clauses: 49
% Final passive clauses: 63
% Main loop iterations started: 56
% 
% >>> Simplifying Inferences
% Duplicate literals: 98
% Fw subsumption resolutions: 50
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 87
% Unit resulting resolution: 2
% Forward superposition: 4
% Backward superposition: 2
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 20
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 27
% SAT solver unit clauses: 4
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

