i: [1,50,0,3,1,1,0,0] t8_zfmisc_1
a: [3,78,0,3,1,50]
s: 78
i: [1,49,0,8,1,1,0,0] t8_zfmisc_1
a: [3,83,0,8,1,49]
s: 83
i: [1,51,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,88,0,2,1,51]
s: 88
i: [1,52,0,2,1,0,0,255] fc1_xboole_0
a: [3,93,0,2,1,52]
s: 93
i: [1,70,0,2,1,0,0,2] rc2_xboole_0
a: [3,103,0,2,1,70]
s: 103
i: [1,71,0,2,1,0,0,2] rc1_xboole_0
a: [3,108,0,2,1,71]
s: 108
i: [1,54,0,3,1,0,0,1] fc2_xboole_0
s: 54
i: [1,64,0,3,1,0,0,255] reflexivity_r1_tarski
s: 64
i: [1,53,0,3,1,0,0,255] d2_xboole_0
a: [3,98,0,3,1,53]
s: 98
r: 88
i: [1,73,0,4,1,0,0,1] d1_tarski
s: 73
i: [1,55,0,5,1,0,0,255] t1_boole
s: 55
i: [1,56,0,5,2,0,0,255] t6_boole
s: 56
i: [1,63,0,5,2,0,0,2] t7_boole
s: 63
i: [1,65,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 65
i: [1,67,0,6,2,0,0,1] fc5_xboole_0
s: 67
i: [1,68,0,6,2,0,0,1] fc4_xboole_0
s: 68
i: [1,69,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 69
i: [1,62,0,7,3,0,0,1] t8_boole
s: 62
d: [2,121,1,4,1,0,40,69,73]
s: 121
i: [1,66,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 66
i: [1,74,0,7,2,0,0,1] d1_tarski
s: 74
d: [2,122,1,5,1,0,44,66,55]
s: 122
i: [1,57,0,9,2,0,0,255] l15_zfmisc_1
s: 57
d: [2,138,1,9,2,1,44,57,83]
a: [3,150,1,5,1,138]
s: 150
i: [1,60,0,14,3,0,0,1] d1_tarski
s: 60
i: [1,61,0,14,3,0,0,1] d1_tarski
s: 61
d: [2,123,1,5,1,0,44,66,55]
s: 123
d: [2,124,1,5,1,0,44,55,66]
s: 124
d: [2,127,1,5,1,0,44,55,66]
s: 127
d: [2,111,1,7,3,0,44,56,56]
s: 111
d: [2,112,1,7,2,0,44,55,56]
s: 112
d: [2,158,2,9,3,0,44,57,112]
d: [2,164,2,6,2,0,21,158,63]
s: 164
d: [2,139,1,7,2,0,44,57,55]
s: 139
d: [2,131,2,7,2,0,44,122,56]
s: 131
d: [2,141,1,9,2,0,44,57,66]
s: 141
d: [2,142,1,9,2,0,44,57,66]
s: 142
d: [2,180,2,9,2,1,44,141,83]
d: [2,187,2,4,1,1,21,180,64]
a: [3,191,2,4,1,187]
s: 191
d: [2,193,3,0,0,2,52,78,191,74]
e: 193
d: [2,194,3,3,1,1,40,191,74]
d: [2,198,3,0,0,2,21,194,78]
e: 198
f: 198,187,49,50
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t8_zfmisc_1
% SZS output start Proof for zfmisc_1__t8_zfmisc_1
1. ! [X0] : ! [X1] : (k1_tarski(X0) = k2_xboole_0(k1_tarski(X0),k1_tarski(X1)) => X0 = X1) [input t8_zfmisc_1]
2. ~! [X0] : ! [X1] : (k1_tarski(X0) = k2_xboole_0(k1_tarski(X0),k1_tarski(X1)) => X0 = X1) [negated conjecture 1]
4. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
5. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
16. ! [X0] : ! [X1] : (r1_tarski(k2_xboole_0(k1_tarski(X0),X1),X1) => r2_hidden(X0,X1)) [input l15_zfmisc_1]
19. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
24. ~! [X0,X1] : (k1_tarski(X0) = k2_xboole_0(k1_tarski(X0),k1_tarski(X1)) => X0 = X1) [flattening 2]
25. ! [X0,X1] : (r1_tarski(k2_xboole_0(k1_tarski(X0),X1),X1) => r2_hidden(X0,X1)) [flattening 16]
26. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 5]
29. ! [X0] : r1_tarski(X0,X0) [rectify 19]
31. ? [X0,X1] : (X0 != X1 & k1_tarski(X0) = k2_xboole_0(k1_tarski(X0),k1_tarski(X1))) [ennf transformation 24]
33. ! [X0,X1] : (r2_hidden(X0,X1) | ~r1_tarski(k2_xboole_0(k1_tarski(X0),X1),X1)) [ennf transformation 25]
39. ? [X0,X1] : (X0 != X1 & k1_tarski(X0) = k2_xboole_0(k1_tarski(X0),k1_tarski(X1))) => (sK0 != sK1 & k1_tarski(sK0) = k2_xboole_0(k1_tarski(sK0),k1_tarski(sK1))) [choice axiom]
40. sK0 != sK1 & k1_tarski(sK0) = k2_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) [skolemisation 31,39]
41. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 26]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 41]
43. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) [choice axiom]
44. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 42,43]
49. k1_tarski(sK0) = k2_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) [cnf transformation 40] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. sK0 != sK1 [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
57. ~r1_tarski(k2_xboole_0(k1_tarski(X0),X1),X1) | r2_hidden(X0,X1) [cnf transformation 33] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
58. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 44] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
64. r1_tarski(X0,X0) [cnf transformation 29] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
66. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
74. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 58] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
76. 1 <=> sK0 = sK1 [avatar definition]
78. sK0 != sK1 <- (~1) [avatar component clause 76] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. ~1 [avatar split clause 50,76]
81. 2 <=> k1_tarski(sK0) = k2_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) [avatar definition]
83. k1_tarski(sK0) = k2_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) <- (2) [avatar component clause 81] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. 2 [avatar split clause 49,81]
141. ~r1_tarski(k2_xboole_0(X3,k1_tarski(X2)),X3) | r2_hidden(X2,X3) [superposition 57,66] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
180. ~r1_tarski(k1_tarski(sK0),k1_tarski(sK0)) | r2_hidden(sK1,k1_tarski(sK0)) <- (2) [superposition 141,83] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
187. r2_hidden(sK1,k1_tarski(sK0)) <- (2) [subsumption resolution 180,64] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
189. 10 <=> r2_hidden(sK1,k1_tarski(sK0)) [avatar definition]
191. r2_hidden(sK1,k1_tarski(sK0)) <- (10) [avatar component clause 189] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
192. 10 | ~2 [avatar split clause 187,81,189]
194. sK0 = sK1 <- (10) [resolution 191,74] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
198. $false <- (~1, 10) [subsumption resolution 194,78] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
199. 1 | ~10 [avatar contradiction clause 198]
200. $false [avatar sat refutation 79,84,192,199] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:44,thDist:-44}
% SZS output end Proof for zfmisc_1__t8_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 106
% Active clauses: 32
% Passive clauses: 44
% Final active clauses: 31
% Final passive clauses: 7
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 1
% Forward subsumptions: 41
% 
% >>> Generating Inferences
% Binary resolution: 10
% Unit resulting resolution: 1
% Forward superposition: 32
% Backward superposition: 26
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 7
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

