i: [1,28,0,2,1,1,0,0] t30_rewrite1
a: [3,58,0,2,1,28]
s: 58
i: [1,31,0,4,1,1,0,0] t30_rewrite1
a: [3,43,0,4,1,31]
s: 43
i: [1,30,0,4,1,1,0,0] t30_rewrite1
a: [3,48,0,4,1,30]
s: 48
i: [1,29,0,4,1,1,0,0] t30_rewrite1
a: [3,53,0,4,1,29]
s: 53
i: [1,32,0,5,2,0,0,2] dt_k2_relat_1
s: 32
i: [1,37,0,5,1,0,0,3] idempotence_k2_xboole_0
s: 37
i: [1,33,0,7,2,0,0,2] involutiveness_k2_relat_1
s: 33
i: [1,38,0,7,1,0,0,3] commutativity_k2_xboole_0
s: 38
i: [1,39,0,8,3,0,0,2] fc3_relat_1
s: 39
i: [1,34,0,13,3,0,0,1] d4_rewrite1
s: 34
i: [1,35,0,13,3,0,0,1] d4_rewrite1
s: 35
i: [1,36,0,14,4,0,0,2] t16_rewrite1
s: 36
d: [2,65,1,17,4,0,44,34,33]
d: [2,66,1,14,3,0,21,65,32]
d: [2,67,1,14,3,0,22,66,38]
s: 67
d: [2,74,2,13,4,0,40,67,35]
d: [2,77,2,11,3,0,18,74]
s: 77
d: [2,69,1,17,4,0,44,35,33]
d: [2,71,1,14,3,0,21,69,32]
d: [2,72,1,14,3,0,22,71,38]
s: 72
d: [2,73,1,25,5,0,40,36,34]
s: 73
d: [2,83,2,13,4,0,40,72,34]
d: [2,85,2,11,3,0,18,83]
s: 85
d: [2,91,2,29,6,0,44,73,33]
d: [2,94,2,24,5,0,21,91,39]
d: [2,95,2,21,4,0,21,94,32]
d: [2,96,2,21,4,0,22,95,38]
d: [2,97,2,21,4,0,22,96,38]
s: 97
d: [2,80,3,14,4,0,44,77,33]
d: [2,81,3,11,3,0,21,80,32]
s: 81
d: [2,90,2,24,6,0,40,73,34]
d: [2,92,2,22,5,0,18,90]
s: 92
d: [2,89,2,25,6,0,40,73,67]
d: [2,93,2,23,5,0,18,89]
s: 93
d: [2,102,3,20,5,0,40,97,34]
d: [2,104,3,18,4,0,18,102]
s: 104
d: [2,75,2,26,5,0,40,67,36]
s: 75
d: [2,140,4,17,5,0,40,104,35]
d: [2,143,4,15,4,0,18,140]
s: 143
d: [2,128,3,26,6,0,44,93,33]
d: [2,134,3,21,5,0,21,128,39]
d: [2,135,3,18,4,0,21,134,32]
d: [2,136,3,18,4,0,22,135,38]
s: 136
d: [2,101,3,21,5,0,40,97,67]
d: [2,105,3,19,4,0,18,101]
s: 105
d: [2,149,5,16,5,0,40,143,85]
d: [2,151,5,14,4,0,18,149]
s: 151
d: [2,165,6,0,0,4,52,58,43,53,48,151]
e: 165
f: 165,28,29,30,31
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rewrite1__t30_rewrite1
% SZS output start Proof for rewrite1__t30_rewrite1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : ((r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X3))) [input t30_rewrite1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : ((r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X3))) [negated conjecture 1]
3. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
4. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r2_rewrite1(X0,X1,X2) <=> r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2))) [input d4_rewrite1]
5. ! [X0] : (v1_relat_1(X0) => v1_relat_1(k2_relat_1(X0))) [input dt_k2_relat_1]
7. ! [X0,X1] : ((v1_relat_1(X1) & v1_relat_1(X0)) => v1_relat_1(k2_xboole_0(X0,X1))) [input fc3_relat_1]
9. ! [X0] : (v1_relat_1(X0) => k2_relat_1(k2_relat_1(X0)) = X0) [input involutiveness_k2_relat_1]
10. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : ((r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2)) => r1_rewrite1(X0,X1,X3))) [input t16_rewrite1]
11. ~! [X0] : (v1_relat_1(X0) => ! [X1,X2,X3] : ((r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X3))) [flattening 2]
12. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r2_rewrite1(X0,X1,X2) <=> r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2))) [flattening 4]
13. ! [X0] : (v1_relat_1(X0) => ! [X1,X2,X3] : ((r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2)) => r1_rewrite1(X0,X1,X3))) [flattening 10]
15. ? [X0] : (? [X1,X2,X3] : (~r2_rewrite1(X0,X1,X3) & (r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2))) & v1_relat_1(X0)) [ennf transformation 11]
16. ? [X0] : (? [X1,X2,X3] : (~r2_rewrite1(X0,X1,X3) & r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2)) & v1_relat_1(X0)) [flattening 15]
17. ! [X0] : (v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 5]
18. ! [X0] : (k2_relat_1(k2_relat_1(X0)) = X0 | ~v1_relat_1(X0)) [ennf transformation 9]
19. ! [X0] : (! [X1,X2] : (r2_rewrite1(X0,X1,X2) <=> r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2)) | ~v1_relat_1(X0)) [ennf transformation 12]
20. ! [X0] : (! [X1,X2,X3] : (r1_rewrite1(X0,X1,X3) | (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [ennf transformation 13]
21. ! [X0] : (! [X1,X2,X3] : (r1_rewrite1(X0,X1,X3) | ~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X2)) | ~v1_relat_1(X0)) [flattening 20]
22. ! [X0,X1] : (v1_relat_1(k2_xboole_0(X0,X1)) | (~v1_relat_1(X1) | ~v1_relat_1(X0))) [ennf transformation 7]
23. ! [X0,X1] : (v1_relat_1(k2_xboole_0(X0,X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0)) [flattening 22]
24. ? [X0] : (? [X1,X2,X3] : (~r2_rewrite1(X0,X1,X3) & r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2)) & v1_relat_1(X0)) => (? [X3,X2,X1] : (~r2_rewrite1(sK0,X1,X3) & r2_rewrite1(sK0,X2,X3) & r2_rewrite1(sK0,X1,X2)) & v1_relat_1(sK0)) [choice axiom]
25. ? [X3,X2,X1] : (~r2_rewrite1(sK0,X1,X3) & r2_rewrite1(sK0,X2,X3) & r2_rewrite1(sK0,X1,X2)) => (~r2_rewrite1(sK0,sK1,sK3) & r2_rewrite1(sK0,sK2,sK3) & r2_rewrite1(sK0,sK1,sK2)) [choice axiom]
26. (~r2_rewrite1(sK0,sK1,sK3) & r2_rewrite1(sK0,sK2,sK3) & r2_rewrite1(sK0,sK1,sK2)) & v1_relat_1(sK0) [skolemisation 16,25,24]
27. ! [X0] : (! [X1,X2] : ((r2_rewrite1(X0,X1,X2) | ~r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2)) & (r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2) | ~r2_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [nnf transformation 19]
28. v1_relat_1(sK0) [cnf transformation 26] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
29. r2_rewrite1(sK0,sK1,sK2) [cnf transformation 26] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
30. r2_rewrite1(sK0,sK2,sK3) [cnf transformation 26] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
31. ~r2_rewrite1(sK0,sK1,sK3) [cnf transformation 26] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
32. v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 17] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
33. k2_relat_1(k2_relat_1(X0)) = X0 | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
34. r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2) | ~r2_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
35. ~r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2) | r2_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
36. ~r1_rewrite1(X0,X2,X3) | r1_rewrite1(X0,X1,X3) | ~r1_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 21] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
38. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
39. v1_relat_1(k2_xboole_0(X0,X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
41. 1 <=> r2_rewrite1(sK0,sK1,sK3) [avatar definition]
43. ~r2_rewrite1(sK0,sK1,sK3) <- (~1) [avatar component clause 41] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
44. ~1 [avatar split clause 31,41]
46. 2 <=> r2_rewrite1(sK0,sK2,sK3) [avatar definition]
48. r2_rewrite1(sK0,sK2,sK3) <- (2) [avatar component clause 46] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
49. 2 [avatar split clause 30,46]
51. 3 <=> r2_rewrite1(sK0,sK1,sK2) [avatar definition]
53. r2_rewrite1(sK0,sK1,sK2) <- (3) [avatar component clause 51] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
54. 3 [avatar split clause 29,51]
56. 4 <=> v1_relat_1(sK0) [avatar definition]
58. v1_relat_1(sK0) <- (4) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
59. 4 [avatar split clause 28,56]
69. ~r1_rewrite1(k2_xboole_0(k2_relat_1(X0),X0),X1,X2) | r2_rewrite1(k2_relat_1(X0),X1,X2) | ~v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [superposition 35,33] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
71. ~r1_rewrite1(k2_xboole_0(k2_relat_1(X0),X0),X1,X2) | r2_rewrite1(k2_relat_1(X0),X1,X2) | ~v1_relat_1(X0) [subsumption resolution 69,32] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
72. ~r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2) | r2_rewrite1(k2_relat_1(X0),X1,X2) | ~v1_relat_1(X0) [forward demodulation 71,38] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
73. ~r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X3) | r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2) | ~v1_relat_1(k2_xboole_0(X0,k2_relat_1(X0))) | ~r2_rewrite1(X0,X3,X2) | ~v1_relat_1(X0) [resolution 36,34] {a:1,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
83. r2_rewrite1(k2_relat_1(X3),X4,X5) | ~v1_relat_1(X3) | ~r2_rewrite1(X3,X4,X5) | ~v1_relat_1(X3) [resolution 72,34] {a:2,w:13,wCS:5200,sine:1,thAx:0,allAx:3,thDist:-3}
85. r2_rewrite1(k2_relat_1(X3),X4,X5) | ~v1_relat_1(X3) | ~r2_rewrite1(X3,X4,X5) [duplicate literal removal 83] {a:2,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
91. ~r1_rewrite1(k2_xboole_0(k2_relat_1(X0),X0),X1,X2) | r1_rewrite1(k2_xboole_0(k2_relat_1(X0),X0),X1,X3) | ~v1_relat_1(k2_xboole_0(k2_relat_1(X0),X0)) | ~r2_rewrite1(k2_relat_1(X0),X2,X3) | ~v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [superposition 73,33] {a:2,w:29,wCS:11600,sine:1,thAx:0,allAx:3,thDist:-3}
94. ~r1_rewrite1(k2_xboole_0(k2_relat_1(X0),X0),X1,X2) | r1_rewrite1(k2_xboole_0(k2_relat_1(X0),X0),X1,X3) | ~r2_rewrite1(k2_relat_1(X0),X2,X3) | ~v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [subsumption resolution 91,39] {a:2,w:24,wCS:9600,sine:1,thAx:0,allAx:3,thDist:-3}
95. ~r1_rewrite1(k2_xboole_0(k2_relat_1(X0),X0),X1,X2) | r1_rewrite1(k2_xboole_0(k2_relat_1(X0),X0),X1,X3) | ~r2_rewrite1(k2_relat_1(X0),X2,X3) | ~v1_relat_1(X0) [subsumption resolution 94,32] {a:2,w:21,wCS:8400,sine:1,thAx:0,allAx:3,thDist:-3}
96. ~r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2) | r1_rewrite1(k2_xboole_0(k2_relat_1(X0),X0),X1,X3) | ~r2_rewrite1(k2_relat_1(X0),X2,X3) | ~v1_relat_1(X0) [forward demodulation 95,38] {a:2,w:21,wCS:8400,sine:1,thAx:0,allAx:3,thDist:-3}
97. ~r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X2) | r1_rewrite1(k2_xboole_0(X0,k2_relat_1(X0)),X1,X3) | ~r2_rewrite1(k2_relat_1(X0),X2,X3) | ~v1_relat_1(X0) [forward demodulation 96,38] {a:2,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
102. r1_rewrite1(k2_xboole_0(X4,k2_relat_1(X4)),X5,X6) | ~r2_rewrite1(k2_relat_1(X4),X7,X6) | ~v1_relat_1(X4) | ~r2_rewrite1(X4,X5,X7) | ~v1_relat_1(X4) [resolution 97,34] {a:3,w:20,wCS:8000,sine:1,thAx:0,allAx:4,thDist:-4}
104. r1_rewrite1(k2_xboole_0(X4,k2_relat_1(X4)),X5,X6) | ~r2_rewrite1(k2_relat_1(X4),X7,X6) | ~v1_relat_1(X4) | ~r2_rewrite1(X4,X5,X7) [duplicate literal removal 102] {a:3,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
140. ~r2_rewrite1(k2_relat_1(X14),X15,X16) | ~v1_relat_1(X14) | ~r2_rewrite1(X14,X17,X15) | r2_rewrite1(X14,X17,X16) | ~v1_relat_1(X14) [resolution 104,35] {a:4,w:17,wCS:6800,sine:1,thAx:0,allAx:5,thDist:-5}
143. ~r2_rewrite1(k2_relat_1(X14),X15,X16) | ~v1_relat_1(X14) | ~r2_rewrite1(X14,X17,X15) | r2_rewrite1(X14,X17,X16) [duplicate literal removal 140] {a:4,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
149. ~v1_relat_1(X0) | ~r2_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X1,X3) | ~v1_relat_1(X0) | ~r2_rewrite1(X0,X2,X3) [resolution 143,85] {a:5,w:16,wCS:6400,sine:1,thAx:0,allAx:8,thDist:-8}
151. ~r2_rewrite1(X0,X2,X3) | ~r2_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X1,X3) | ~v1_relat_1(X0) [duplicate literal removal 149] {a:5,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:8,thDist:-8}
165. $false <- (~1, 2, 3, 4) [unit resulting resolution 58,43,53,48,151] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:20,thDist:-20}
169. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 165]
173. $false [avatar sat refutation 44,49,54,59,169] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:36,thDist:-36}
% SZS output end Proof for rewrite1__t30_rewrite1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 128
% Active clauses: 24
% Passive clauses: 51
% Final active clauses: 24
% Final passive clauses: 24
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Duplicate literals: 31
% Fw subsumption resolutions: 18
% Fw demodulations: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 35
% Unit resulting resolution: 1
% Forward superposition: 16
% Backward superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 7
% SAT solver unit clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

