i: [1,53,0,2,1,1,0,0] t11_tsep_1
a: [3,103,0,2,1,53]
s: 103
i: [1,52,0,2,1,1,0,0] t11_tsep_1
a: [3,108,0,2,1,52]
s: 108
i: [1,59,0,9,3,1,0,0] t11_tsep_1
a: [3,77,0,3,1,59]
s: 77
a: [3,80,0,3,1,59]
s: 80
a: [3,84,0,3,1,59]
s: 84
i: [1,56,0,4,1,1,0,0] t11_tsep_1
a: [3,92,0,4,1,56]
s: 92
i: [1,55,0,5,1,1,0,0] t11_tsep_1
a: [3,97,0,5,1,55]
s: 97
i: [1,72,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,113,0,2,1,72]
s: 113
i: [1,71,0,3,1,0,0,255] reflexivity_r1_tarski
s: 71
i: [1,70,0,4,1,0,0,1] existence_m1_subset_1
s: 70
i: [1,61,0,6,2,0,0,1] existence_m1_pre_topc
s: 61
i: [1,60,0,7,3,0,0,1] dt_m1_pre_topc
s: 60
i: [1,69,0,7,2,0,0,255] t3_subset
s: 69
d: [2,115,1,4,2,1,40,60,80]
d: [2,118,1,2,1,2,21,115,103]
a: [3,122,1,2,1,118]
s: 122
i: [1,68,0,8,3,0,0,1] rc6_pre_topc
s: 68
i: [1,62,0,9,4,0,0,1] cc1_pre_topc
s: 62
d: [2,124,1,6,3,1,40,62,80]
d: [2,127,1,4,2,2,21,124,103]
d: [2,128,1,2,1,3,21,127,108]
a: [3,132,1,2,1,128]
s: 132
i: [1,67,0,10,3,0,0,1] rc6_pre_topc
s: 67
d: [2,134,1,9,3,1,44,67,92]
d: [2,135,1,7,2,2,21,134,122]
d: [2,136,1,5,1,3,21,135,132]
a: [3,140,1,5,1,136]
s: 140
i: [1,66,0,15,5,0,0,1] d11_borsuk_1
s: 66
i: [1,65,0,16,5,0,0,1] d11_borsuk_1
s: 65
d: [2,116,1,7,3,0,40,60,61]
d: [2,117,1,5,2,0,18,116]
s: 117
i: [1,64,0,17,5,0,0,1] d11_borsuk_1
s: 64
d: [2,145,1,16,5,1,44,64,92]
d: [2,147,1,14,4,2,21,145,122]
d: [2,148,1,12,3,3,21,147,132]
s: 148
i: [1,73,0,20,6,0,0,1] d11_borsuk_1
s: 73
d: [2,155,1,19,6,1,44,73,92]
d: [2,156,1,17,5,2,21,155,122]
d: [2,157,1,15,4,3,21,156,132]
s: 157
d: [2,159,2,13,4,3,44,157,92]
a: [3,163,2,3,1,159]
s: 163
d: [2,154,1,18,6,1,44,73,92]
s: 154
d: [2,149,2,21,7,3,44,148,65]
d: [2,150,2,15,5,3,18,149]
d: [2,151,2,13,4,3,21,150,122]
d: [2,152,2,11,3,3,21,151,132]
s: 152
d: [2,125,1,9,4,0,40,62,61]
d: [2,126,1,7,3,0,18,125]
s: 126
d: [2,142,1,24,9,0,44,66,65]
d: [2,143,1,14,5,0,18,142]
s: 143
d: [2,183,2,13,5,1,44,143,92]
s: 183
d: [2,184,3,0,0,6,52,108,103,80,77,84,183]
e: 184
d: [2,185,3,10,4,2,40,183,84]
d: [2,187,3,7,3,3,21,185,77]
d: [2,188,3,4,2,4,21,187,80]
d: [2,189,3,2,1,5,21,188,103]
d: [2,190,3,0,0,6,21,189,108]
e: 190
r: 77
r: 84
i: [1,57,0,6,2,1,0,0] t11_tsep_1
i: [1,54,0,3,1,1,0,0] t11_tsep_1
d: [2,177,2,13,5,2,40,154,97]
d: [2,192,2,10,4,3,21,177,80]
d: [2,193,2,8,3,4,21,192,103]
d: [2,194,2,6,2,5,21,193,108]
f: 194,190,52,53,54,55,56,57,59
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tsep_1__t11_tsep_1
% SZS output start Proof for tsep_1__t11_tsep_1
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_pre_topc(X1,X0) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) => (u1_struct_0(X1) = X2 => ((m1_pre_topc(X1,X0) & v1_borsuk_1(X1,X0)) <=> v4_pre_topc(X2,X0)))))) [input t11_tsep_1]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_pre_topc(X1,X0) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) => (u1_struct_0(X1) = X2 => ((m1_pre_topc(X1,X0) & v1_borsuk_1(X1,X0)) <=> v4_pre_topc(X2,X0)))))) [negated conjecture 1]
4. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_pre_topc(X1,X0) => (v1_borsuk_1(X1,X0) <=> ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) => (u1_struct_0(X1) = X2 => v4_pre_topc(X2,X0)))))) [input d11_borsuk_1]
23. ? [X0] : (? [X1] : (? [X2] : ((((m1_pre_topc(X1,X0) & v1_borsuk_1(X1,X0)) <~> v4_pre_topc(X2,X0)) & u1_struct_0(X1) = X2) & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) & m1_pre_topc(X1,X0)) & (l1_pre_topc(X0) & v2_pre_topc(X0))) [ennf transformation 2]
24. ? [X0] : (? [X1] : (? [X2] : (((m1_pre_topc(X1,X0) & v1_borsuk_1(X1,X0)) <~> v4_pre_topc(X2,X0)) & u1_struct_0(X1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) & m1_pre_topc(X1,X0)) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 23]
29. ! [X0] : (! [X1] : ((v1_borsuk_1(X1,X0) <=> ! [X2] : ((v4_pre_topc(X2,X0) | u1_struct_0(X1) != X2) | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))))) | ~m1_pre_topc(X1,X0)) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 4]
30. ! [X0] : (! [X1] : ((v1_borsuk_1(X1,X0) <=> ! [X2] : (v4_pre_topc(X2,X0) | u1_struct_0(X1) != X2 | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))))) | ~m1_pre_topc(X1,X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 29]
34. ? [X0] : (? [X1] : (? [X2] : (((~v4_pre_topc(X2,X0) | (~m1_pre_topc(X1,X0) | ~v1_borsuk_1(X1,X0))) & (v4_pre_topc(X2,X0) | (m1_pre_topc(X1,X0) & v1_borsuk_1(X1,X0)))) & u1_struct_0(X1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) & m1_pre_topc(X1,X0)) & l1_pre_topc(X0) & v2_pre_topc(X0)) [nnf transformation 24]
35. ? [X0] : (? [X1] : (? [X2] : ((~v4_pre_topc(X2,X0) | ~m1_pre_topc(X1,X0) | ~v1_borsuk_1(X1,X0)) & (v4_pre_topc(X2,X0) | (m1_pre_topc(X1,X0) & v1_borsuk_1(X1,X0))) & u1_struct_0(X1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) & m1_pre_topc(X1,X0)) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 34]
36. ? [X0] : (? [X1] : (? [X2] : ((~v4_pre_topc(X2,X0) | ~m1_pre_topc(X1,X0) | ~v1_borsuk_1(X1,X0)) & (v4_pre_topc(X2,X0) | (m1_pre_topc(X1,X0) & v1_borsuk_1(X1,X0))) & u1_struct_0(X1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) & m1_pre_topc(X1,X0)) & l1_pre_topc(X0) & v2_pre_topc(X0)) => (? [X1] : (? [X2] : ((~v4_pre_topc(X2,sK0) | ~m1_pre_topc(X1,sK0) | ~v1_borsuk_1(X1,sK0)) & (v4_pre_topc(X2,sK0) | (m1_pre_topc(X1,sK0) & v1_borsuk_1(X1,sK0))) & u1_struct_0(X1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_pre_topc(X1,sK0)) & l1_pre_topc(sK0) & v2_pre_topc(sK0)) [choice axiom]
37. ? [X1] : (? [X2] : ((~v4_pre_topc(X2,sK0) | ~m1_pre_topc(X1,sK0) | ~v1_borsuk_1(X1,sK0)) & (v4_pre_topc(X2,sK0) | (m1_pre_topc(X1,sK0) & v1_borsuk_1(X1,sK0))) & u1_struct_0(X1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_pre_topc(X1,sK0)) => (? [X2] : ((~v4_pre_topc(X2,sK0) | ~m1_pre_topc(sK1,sK0) | ~v1_borsuk_1(sK1,sK0)) & (v4_pre_topc(X2,sK0) | (m1_pre_topc(sK1,sK0) & v1_borsuk_1(sK1,sK0))) & u1_struct_0(sK1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_pre_topc(sK1,sK0)) [choice axiom]
38. ? [X2] : ((~v4_pre_topc(X2,sK0) | ~m1_pre_topc(sK1,sK0) | ~v1_borsuk_1(sK1,sK0)) & (v4_pre_topc(X2,sK0) | (m1_pre_topc(sK1,sK0) & v1_borsuk_1(sK1,sK0))) & u1_struct_0(sK1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(sK0)))) => ((~v4_pre_topc(sK2,sK0) | ~m1_pre_topc(sK1,sK0) | ~v1_borsuk_1(sK1,sK0)) & (v4_pre_topc(sK2,sK0) | (m1_pre_topc(sK1,sK0) & v1_borsuk_1(sK1,sK0))) & u1_struct_0(sK1) = sK2 & m1_subset_1(sK2,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
39. (((~v4_pre_topc(sK2,sK0) | ~m1_pre_topc(sK1,sK0) | ~v1_borsuk_1(sK1,sK0)) & (v4_pre_topc(sK2,sK0) | (m1_pre_topc(sK1,sK0) & v1_borsuk_1(sK1,sK0))) & u1_struct_0(sK1) = sK2 & m1_subset_1(sK2,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_pre_topc(sK1,sK0)) & l1_pre_topc(sK0) & v2_pre_topc(sK0) [skolemisation 35,38,37,36]
42. ! [X0] : (! [X1] : (((v1_borsuk_1(X1,X0) | ? [X2] : (~v4_pre_topc(X2,X0) & u1_struct_0(X1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))))) & (! [X2] : (v4_pre_topc(X2,X0) | u1_struct_0(X1) != X2 | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) | ~v1_borsuk_1(X1,X0))) | ~m1_pre_topc(X1,X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [nnf transformation 30]
43. ! [X0] : (! [X1] : (((v1_borsuk_1(X1,X0) | ? [X2] : (~v4_pre_topc(X2,X0) & u1_struct_0(X1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))))) & (! [X3] : (v4_pre_topc(X3,X0) | u1_struct_0(X1) != X3 | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0)))) | ~v1_borsuk_1(X1,X0))) | ~m1_pre_topc(X1,X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [rectify 42]
44. ! [X1,X0] : (? [X2] : (~v4_pre_topc(X2,X0) & u1_struct_0(X1) = X2 & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) => (~v4_pre_topc(sK4(X0,X1),X0) & u1_struct_0(X1) = sK4(X0,X1) & m1_subset_1(sK4(X0,X1),k1_zfmisc_1(u1_struct_0(X0))))) [choice axiom]
45. ! [X0] : (! [X1] : (((v1_borsuk_1(X1,X0) | (~v4_pre_topc(sK4(X0,X1),X0) & u1_struct_0(X1) = sK4(X0,X1) & m1_subset_1(sK4(X0,X1),k1_zfmisc_1(u1_struct_0(X0))))) & (! [X3] : (v4_pre_topc(X3,X0) | u1_struct_0(X1) != X3 | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0)))) | ~v1_borsuk_1(X1,X0))) | ~m1_pre_topc(X1,X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [skolemisation 43,44]
52. v2_pre_topc(sK0) [cnf transformation 39] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
53. l1_pre_topc(sK0) [cnf transformation 39] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
54. m1_pre_topc(sK1,sK0) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
55. m1_subset_1(sK2,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 39] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
56. u1_struct_0(sK1) = sK2 [cnf transformation 39] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
57. v4_pre_topc(sK2,sK0) | v1_borsuk_1(sK1,sK0) [cnf transformation 39] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
59. ~v4_pre_topc(sK2,sK0) | ~m1_pre_topc(sK1,sK0) | ~v1_borsuk_1(sK1,sK0) [cnf transformation 39] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
63. v4_pre_topc(X3,X0) | u1_struct_0(X1) != X3 | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0))) | ~v1_borsuk_1(X1,X0) | ~m1_pre_topc(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 45] {a:0,w:22,wCS:8800,sine:1,thAx:0,allAx:2,thDist:-2}
65. u1_struct_0(X1) = sK4(X0,X1) | v1_borsuk_1(X1,X0) | ~m1_pre_topc(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 45] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
66. ~v4_pre_topc(sK4(X0,X1),X0) | v1_borsuk_1(X1,X0) | ~m1_pre_topc(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 45] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
73. ~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X0))) | v4_pre_topc(u1_struct_0(X1),X0) | ~v1_borsuk_1(X1,X0) | ~m1_pre_topc(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [equality resolution 63] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
75. 1 <=> v1_borsuk_1(sK1,sK0) [avatar definition]
77. ~v1_borsuk_1(sK1,sK0) <- (~1) [avatar component clause 75] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
79. 2 <=> m1_pre_topc(sK1,sK0) [avatar definition]
80. m1_pre_topc(sK1,sK0) <- (2) [avatar component clause 79] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
83. 3 <=> v4_pre_topc(sK2,sK0) [avatar definition]
84. v4_pre_topc(sK2,sK0) <- (3) [avatar component clause 83] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
86. ~1 | ~2 | ~3 [avatar split clause 59,83,79,75]
88. 1 | 3 [avatar split clause 57,83,75]
90. 4 <=> u1_struct_0(sK1) = sK2 [avatar definition]
92. u1_struct_0(sK1) = sK2 <- (4) [avatar component clause 90] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
93. 4 [avatar split clause 56,90]
95. 5 <=> m1_subset_1(sK2,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
97. m1_subset_1(sK2,k1_zfmisc_1(u1_struct_0(sK0))) <- (5) [avatar component clause 95] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
98. 5 [avatar split clause 55,95]
99. 2 [avatar split clause 54,79]
101. 6 <=> l1_pre_topc(sK0) [avatar definition]
103. l1_pre_topc(sK0) <- (6) [avatar component clause 101] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
104. 6 [avatar split clause 53,101]
106. 7 <=> v2_pre_topc(sK0) [avatar definition]
108. v2_pre_topc(sK0) <- (7) [avatar component clause 106] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
109. 7 [avatar split clause 52,106]
142. ~v4_pre_topc(u1_struct_0(X1),X0) | v1_borsuk_1(X1,X0) | ~m1_pre_topc(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v1_borsuk_1(X1,X0) | ~m1_pre_topc(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [superposition 66,65] {a:1,w:24,wCS:9600,sine:1,thAx:0,allAx:4,thDist:-4}
143. ~v4_pre_topc(u1_struct_0(X1),X0) | v1_borsuk_1(X1,X0) | ~m1_pre_topc(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [duplicate literal removal 142] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
154. ~m1_subset_1(sK2,k1_zfmisc_1(u1_struct_0(X0))) | v4_pre_topc(sK2,X0) | ~v1_borsuk_1(sK1,X0) | ~m1_pre_topc(sK1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) <- (4) [superposition 73,92] {a:1,w:18,wCS:1800,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
177. v4_pre_topc(sK2,sK0) | ~v1_borsuk_1(sK1,sK0) | ~m1_pre_topc(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (4, 5) [resolution 154,97] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
183. ~v4_pre_topc(sK2,X0) | v1_borsuk_1(sK1,X0) | ~m1_pre_topc(sK1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) <- (4) [superposition 143,92] {a:2,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
185. v1_borsuk_1(sK1,sK0) | ~m1_pre_topc(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (3, 4) [resolution 183,84] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
187. ~m1_pre_topc(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~1, 3, 4) [subsumption resolution 185,77] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
188. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~1, 2, 3, 4) [subsumption resolution 187,80] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
189. ~v2_pre_topc(sK0) <- (~1, 2, 3, 4, 6) [subsumption resolution 188,103] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
190. $false <- (~1, 2, 3, 4, 6, 7) [subsumption resolution 189,108] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
191. 1 | ~2 | ~3 | ~4 | ~6 | ~7 [avatar contradiction clause 190]
192. v4_pre_topc(sK2,sK0) | ~v1_borsuk_1(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (2, 4, 5) [subsumption resolution 177,80] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
193. v4_pre_topc(sK2,sK0) | ~v1_borsuk_1(sK1,sK0) | ~v2_pre_topc(sK0) <- (2, 4, 5, 6) [subsumption resolution 192,103] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
194. v4_pre_topc(sK2,sK0) | ~v1_borsuk_1(sK1,sK0) <- (2, 4, 5, 6, 7) [subsumption resolution 193,108] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
195. ~1 | 3 | ~2 | ~4 | ~5 | ~6 | ~7 [avatar split clause 194,106,101,95,90,79,83,75]
196. $false [avatar sat refutation 86,88,93,98,99,104,109,191,195] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:148,thDist:-148}
% SZS output end Proof for tsep_1__t11_tsep_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 82
% Active clauses: 32
% Passive clauses: 38
% Final active clauses: 30
% Final passive clauses: 6
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Duplicate literals: 13
% Fw subsumption resolutions: 18
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 10
% Unit resulting resolution: 1
% Forward superposition: 10
% Backward superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 13
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 6
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

