i: [1,28,0,2,1,1,0,0] t229_xreal_1
a: [3,52,0,2,1,28]
s: 52
i: [1,27,0,2,1,1,0,0] t229_xreal_1
a: [3,57,0,2,1,27]
s: 57
i: [1,26,0,2,1,1,0,0] t229_xreal_1
a: [3,62,0,2,1,26]
s: 62
i: [1,31,0,3,1,1,0,0] t229_xreal_1
a: [3,42,0,3,1,31]
s: 42
i: [1,30,0,3,1,1,0,0] t229_xreal_1
a: [3,47,0,3,1,30]
s: 47
i: [1,38,0,2,1,0,0,1] rc1_xxreal_0
a: [3,74,0,2,1,38]
s: 74
i: [1,29,0,11,4,1,0,0] t229_xreal_1
s: 29
i: [1,36,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,69,0,5,2,36]
s: 69
d: [2,76,1,8,3,1,40,29,42]
d: [2,78,1,6,2,2,21,76,52]
a: [3,82,1,3,1,78]
s: 82
i: [1,37,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 37
i: [1,32,0,11,4,0,0,1] t227_xreal_1
s: 32
d: [2,77,1,8,3,1,40,29,47]
d: [2,88,1,6,2,2,21,77,62]
a: [3,92,1,3,1,88]
s: 92
i: [1,33,0,12,4,0,0,1] t227_xreal_1
s: 33
d: [2,98,1,7,3,1,40,37,42]
d: [2,102,1,5,2,2,21,98,57]
d: [2,103,1,3,1,3,21,102,52]
a: [3,107,1,3,1,103]
s: 107
i: [1,34,0,12,4,0,0,1] t227_xreal_1
s: 34
i: [1,35,0,15,6,0,0,1] t2_xxreal_0
s: 35
d: [2,99,1,7,3,1,40,37,47]
d: [2,109,1,5,2,2,21,99,57]
d: [2,110,1,3,1,3,21,109,62]
a: [3,114,1,3,1,110]
s: 114
d: [2,124,1,20,8,0,40,35,29]
d: [2,133,1,18,7,0,18,124]
d: [2,134,1,16,6,1,21,133,57]
s: 134
d: [2,127,2,12,5,1,40,35,82]
d: [2,135,2,10,4,2,21,127,62]
d: [2,136,2,8,3,3,21,135,52]
s: 136
d: [2,119,1,21,6,0,40,34,29]
d: [2,122,1,19,5,1,21,119,57]
s: 122
d: [2,157,1,14,4,4,21,122,136]
s: 157
d: [2,158,2,16,6,4,40,157,33]
d: [2,161,2,11,4,4,18,158]
d: [2,162,2,9,3,4,21,161,62]
d: [2,163,2,6,2,5,21,162,47]
d: [2,164,2,4,1,5,21,163,57]
a: [3,168,2,4,1,164]
s: 168
d: [2,171,3,0,0,4,52,62,57,47,168,32]
e: 171
d: [2,172,3,7,3,1,40,168,32]
d: [2,174,3,4,2,2,21,172,47]
d: [2,175,3,2,1,3,21,174,57]
d: [2,176,3,0,0,4,21,175,62]
e: 176
r: 157
r: 136
r: 82
r: 168
a: [3,81,1,3,1,78]
s: 81
a: [3,86,1,3,1,78]
s: 86
a: [3,96,1,3,1,88]
s: 96
s: 122
a: [3,167,2,4,1,164]
s: 167
d: [2,116,1,18,6,0,40,33,37]
d: [2,117,1,16,5,0,18,116]
d: [2,118,1,12,4,0,21,117,32]
s: 118
d: [2,187,2,21,7,1,40,122,33]
d: [2,190,2,16,5,1,18,187]
d: [2,191,2,14,4,2,21,190,52]
d: [2,192,2,11,3,3,21,191,42]
d: [2,193,2,9,2,3,21,192,57]
a: [3,197,2,4,1,193]
s: 197
d: [2,209,3,0,0,4,52,52,57,42,197,32]
e: 209
d: [2,210,3,7,3,1,40,197,32]
d: [2,212,3,4,2,2,21,210,42]
d: [2,213,3,2,1,3,21,212,57]
d: [2,214,3,0,0,4,21,213,52]
e: 214
r: 197
d: [2,120,1,18,6,0,40,34,37]
d: [2,121,1,16,5,0,18,120]
d: [2,123,1,12,4,0,21,121,32]
s: 123
d: [2,126,1,19,8,0,40,35,37]
d: [2,131,1,15,6,0,18,126]
s: 131
a: [3,196,2,4,1,193]
s: 196
a: [3,201,2,5,1,193]
s: 201
d: [2,129,2,12,5,1,40,35,107]
d: [2,137,2,10,4,2,21,129,57]
d: [2,138,2,8,3,3,21,137,52]
s: 138
d: [2,139,2,12,5,1,40,114,35]
d: [2,140,2,10,4,2,21,139,57]
d: [2,141,2,8,3,3,21,140,62]
s: 141
d: [2,184,2,12,5,1,40,96,35]
d: [2,185,2,10,4,2,21,184,52]
d: [2,186,2,8,3,3,21,185,62]
s: 186
d: [2,271,3,9,2,4,40,141,201]
d: [2,278,3,5,1,5,21,271,196]
a: [3,282,3,5,1,278]
s: 282
d: [2,228,2,12,5,1,40,131,107]
d: [2,244,2,10,4,2,21,228,57]
d: [2,245,2,8,3,3,21,244,52]
s: 245
d: [2,287,3,9,2,4,40,186,201]
d: [2,294,3,5,1,5,21,287,196]
a: [3,298,3,5,1,294]
s: 298
d: [2,325,4,0,0,4,52,52,57,42,298,33]
e: 325
d: [2,326,4,7,3,1,40,298,33]
d: [2,331,4,4,2,2,21,326,42]
d: [2,332,4,2,1,3,21,331,57]
d: [2,333,4,0,0,4,21,332,52]
e: 333
d: [2,178,2,7,3,1,40,81,37]
d: [2,179,2,5,2,2,21,178,52]
d: [2,180,2,3,1,3,21,179,62]
f: 333,294,214,193,180,176,164,26,27,28,30,31
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t229_xreal_1
% SZS output start Proof for xreal_1__t229_xreal_1
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (! [X3] : (v1_xxreal_0(X3) => ~(~r1_xxreal_0(X3,X2) & ~r1_xxreal_0(X1,X3) & ~r1_xxreal_0(X3,X0))) => (r1_xxreal_0(X1,X2) | r1_xxreal_0(X1,X0)))))) [input t229_xreal_1]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (! [X3] : (v1_xxreal_0(X3) => ~(~r1_xxreal_0(X3,X2) & ~r1_xxreal_0(X1,X3) & ~r1_xxreal_0(X3,X0))) => (r1_xxreal_0(X1,X2) | r1_xxreal_0(X1,X0)))))) [negated conjecture 1]
3. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
6. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ~(! [X2] : (v1_xxreal_0(X2) => ~(~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X2,X0))) & ~r1_xxreal_0(X1,X0)))) [input t227_xreal_1]
7. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input t2_xxreal_0]
8. ? [X0] : (? [X1] : (? [X2] : (((~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X1,X0)) & ! [X3] : ((r1_xxreal_0(X3,X2) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,X0)) | ~v1_xxreal_0(X3))) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
9. ? [X0] : (? [X1] : (? [X2] : (~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X1,X0) & ! [X3] : (r1_xxreal_0(X3,X2) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,X0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 8]
10. ! [X0] : (! [X1] : ((? [X2] : ((~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X2,X0)) & v1_xxreal_0(X2)) | r1_xxreal_0(X1,X0)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 6]
11. ! [X0] : (! [X1] : (? [X2] : (~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X2,X0) & v1_xxreal_0(X2)) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 10]
12. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 7]
13. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 12]
16. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 3]
17. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 16]
18. ? [X0] : (? [X1] : (? [X2] : (~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X1,X0) & ! [X3] : (r1_xxreal_0(X3,X2) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,X0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X1,sK0) & ! [X3] : (r1_xxreal_0(X3,X2) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
19. ? [X1] : (? [X2] : (~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X1,sK0) & ! [X3] : (r1_xxreal_0(X3,X2) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (~r1_xxreal_0(sK1,X2) & ~r1_xxreal_0(sK1,sK0) & ! [X3] : (r1_xxreal_0(X3,X2) | r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
20. ? [X2] : (~r1_xxreal_0(sK1,X2) & ~r1_xxreal_0(sK1,sK0) & ! [X3] : (r1_xxreal_0(X3,X2) | r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) => (~r1_xxreal_0(sK1,sK2) & ~r1_xxreal_0(sK1,sK0) & ! [X3] : (r1_xxreal_0(X3,sK2) | r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(sK2)) [choice axiom]
21. ((~r1_xxreal_0(sK1,sK2) & ~r1_xxreal_0(sK1,sK0) & ! [X3] : (r1_xxreal_0(X3,sK2) | r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 9,20,19,18]
22. ! [X1,X0] : (? [X2] : (~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X2,X0) & v1_xxreal_0(X2)) => (~r1_xxreal_0(X1,sK3(X0,X1)) & ~r1_xxreal_0(sK3(X0,X1),X0) & v1_xxreal_0(sK3(X0,X1)))) [choice axiom]
23. ! [X0] : (! [X1] : ((~r1_xxreal_0(X1,sK3(X0,X1)) & ~r1_xxreal_0(sK3(X0,X1),X0) & v1_xxreal_0(sK3(X0,X1))) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [skolemisation 11,22]
26. v1_xxreal_0(sK0) [cnf transformation 21] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
27. v1_xxreal_0(sK1) [cnf transformation 21] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
28. v1_xxreal_0(sK2) [cnf transformation 21] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
29. r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK2) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3) [cnf transformation 21] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
30. ~r1_xxreal_0(sK1,sK0) [cnf transformation 21] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
31. ~r1_xxreal_0(sK1,sK2) [cnf transformation 21] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
32. v1_xxreal_0(sK3(X0,X1)) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 23] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
33. ~r1_xxreal_0(sK3(X0,X1),X0) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 23] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
34. ~r1_xxreal_0(X1,sK3(X0,X1)) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 23] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
35. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 13] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 17] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
40. 1 <=> r1_xxreal_0(sK1,sK2) [avatar definition]
42. ~r1_xxreal_0(sK1,sK2) <- (~1) [avatar component clause 40] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
43. ~1 [avatar split clause 31,40]
45. 2 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
47. ~r1_xxreal_0(sK1,sK0) <- (~2) [avatar component clause 45] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
48. ~2 [avatar split clause 30,45]
50. 3 <=> v1_xxreal_0(sK2) [avatar definition]
52. v1_xxreal_0(sK2) <- (3) [avatar component clause 50] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
53. 3 [avatar split clause 28,50]
55. 4 <=> v1_xxreal_0(sK1) [avatar definition]
57. v1_xxreal_0(sK1) <- (4) [avatar component clause 55] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
58. 4 [avatar split clause 27,55]
60. 5 <=> v1_xxreal_0(sK0) [avatar definition]
62. v1_xxreal_0(sK0) <- (5) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
63. 5 [avatar split clause 26,60]
80. 9 <=> r1_xxreal_0(sK2,sK0) [avatar definition]
81. ~r1_xxreal_0(sK2,sK0) <- (~9) [avatar component clause 80] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
82. r1_xxreal_0(sK2,sK0) <- (9) [avatar component clause 80] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
94. 12 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
96. r1_xxreal_0(sK0,sK2) <- (12) [avatar component clause 94] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
119. r1_xxreal_0(sK1,X0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(X0) | r1_xxreal_0(sK3(X0,sK1),sK2) | r1_xxreal_0(sK3(X0,sK1),sK0) | ~v1_xxreal_0(sK3(X0,sK1)) [resolution 34,29] {a:1,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
122. r1_xxreal_0(sK3(X0,sK1),sK2) | ~v1_xxreal_0(X0) | r1_xxreal_0(sK1,X0) | r1_xxreal_0(sK3(X0,sK1),sK0) | ~v1_xxreal_0(sK3(X0,sK1)) <- (4) [subsumption resolution 119,57] {a:1,w:19,wCS:1900,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
127. r1_xxreal_0(X7,sK0) | ~r1_xxreal_0(X7,sK2) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X7) <- (9) [resolution 35,82] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
135. r1_xxreal_0(X7,sK0) | ~r1_xxreal_0(X7,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X7) <- (5, 9) [subsumption resolution 127,62] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
136. ~r1_xxreal_0(X7,sK2) | r1_xxreal_0(X7,sK0) | ~v1_xxreal_0(X7) <- (3, 5, 9) [subsumption resolution 135,52] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
157. r1_xxreal_0(sK3(X0,sK1),sK0) | ~v1_xxreal_0(X0) | r1_xxreal_0(sK1,X0) | ~v1_xxreal_0(sK3(X0,sK1)) <- (3, 4, 5, 9) [subsumption resolution 122,136] {a:1,w:14,wCS:1400,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
158. ~v1_xxreal_0(sK0) | r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK3(sK0,sK1)) | r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3, 4, 5, 9) [resolution 157,33] {a:2,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
161. ~v1_xxreal_0(sK0) | r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK3(sK0,sK1)) | ~v1_xxreal_0(sK1) <- (3, 4, 5, 9) [duplicate literal removal 158] {a:2,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
162. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK3(sK0,sK1)) | ~v1_xxreal_0(sK1) <- (3, 4, 5, 9) [subsumption resolution 161,62] {a:2,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
163. ~v1_xxreal_0(sK3(sK0,sK1)) | ~v1_xxreal_0(sK1) <- (~2, 3, 4, 5, 9) [subsumption resolution 162,47] {a:2,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
164. ~v1_xxreal_0(sK3(sK0,sK1)) <- (~2, 3, 4, 5, 9) [subsumption resolution 163,57] {a:2,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
166. 15 <=> v1_xxreal_0(sK3(sK0,sK1)) [avatar definition]
168. ~v1_xxreal_0(sK3(sK0,sK1)) <- (~15) [avatar component clause 166] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
169. ~15 | 2 | ~3 | ~4 | ~5 | ~9 [avatar split clause 164,80,60,55,50,45,166]
172. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~15) [resolution 168,32] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
174. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, ~15) [subsumption resolution 172,47] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
175. ~v1_xxreal_0(sK0) <- (~2, 4, ~15) [subsumption resolution 174,57] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
176. $false <- (~2, 4, 5, ~15) [subsumption resolution 175,62] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
177. 2 | ~4 | ~5 | 15 [avatar contradiction clause 176]
178. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (~9) [resolution 81,37] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
179. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK0) <- (3, ~9) [subsumption resolution 178,52] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
180. r1_xxreal_0(sK0,sK2) <- (3, 5, ~9) [subsumption resolution 179,62] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
181. 12 | ~3 | ~5 | 9 [avatar split clause 180,80,60,50,94]
184. r1_xxreal_0(X0,sK2) | ~r1_xxreal_0(X0,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(X0) <- (12) [resolution 96,35] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
185. r1_xxreal_0(X0,sK2) | ~r1_xxreal_0(X0,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(X0) <- (3, 12) [subsumption resolution 184,52] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
186. ~r1_xxreal_0(X0,sK0) | r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X0) <- (3, 5, 12) [subsumption resolution 185,62] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
187. ~v1_xxreal_0(sK2) | r1_xxreal_0(sK1,sK2) | r1_xxreal_0(sK3(sK2,sK1),sK0) | ~v1_xxreal_0(sK3(sK2,sK1)) | r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (4) [resolution 122,33] {a:2,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
190. ~v1_xxreal_0(sK2) | r1_xxreal_0(sK1,sK2) | r1_xxreal_0(sK3(sK2,sK1),sK0) | ~v1_xxreal_0(sK3(sK2,sK1)) | ~v1_xxreal_0(sK1) <- (4) [duplicate literal removal 187] {a:2,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
191. r1_xxreal_0(sK1,sK2) | r1_xxreal_0(sK3(sK2,sK1),sK0) | ~v1_xxreal_0(sK3(sK2,sK1)) | ~v1_xxreal_0(sK1) <- (3, 4) [subsumption resolution 190,52] {a:2,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
192. r1_xxreal_0(sK3(sK2,sK1),sK0) | ~v1_xxreal_0(sK3(sK2,sK1)) | ~v1_xxreal_0(sK1) <- (~1, 3, 4) [subsumption resolution 191,42] {a:2,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
193. r1_xxreal_0(sK3(sK2,sK1),sK0) | ~v1_xxreal_0(sK3(sK2,sK1)) <- (~1, 3, 4) [subsumption resolution 192,57] {a:2,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
195. 16 <=> v1_xxreal_0(sK3(sK2,sK1)) [avatar definition]
196. v1_xxreal_0(sK3(sK2,sK1)) <- (16) [avatar component clause 195] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
197. ~v1_xxreal_0(sK3(sK2,sK1)) <- (~16) [avatar component clause 195] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
199. 17 <=> r1_xxreal_0(sK3(sK2,sK1),sK0) [avatar definition]
201. r1_xxreal_0(sK3(sK2,sK1),sK0) <- (17) [avatar component clause 199] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
202. ~16 | 17 | 1 | ~3 | ~4 [avatar split clause 193,55,50,40,199,195]
210. r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (~16) [resolution 197,32] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
212. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (~1, ~16) [subsumption resolution 210,42] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
213. ~v1_xxreal_0(sK2) <- (~1, 4, ~16) [subsumption resolution 212,57] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
214. $false <- (~1, 3, 4, ~16) [subsumption resolution 213,52] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
215. 1 | ~3 | ~4 | 16 [avatar contradiction clause 214]
287. r1_xxreal_0(sK3(sK2,sK1),sK2) | ~v1_xxreal_0(sK3(sK2,sK1)) <- (3, 5, 12, 17) [resolution 186,201] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:17,thDist:-17}
294. r1_xxreal_0(sK3(sK2,sK1),sK2) <- (3, 5, 12, 16, 17) [subsumption resolution 287,196] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:17,thDist:-17}
296. 19 <=> r1_xxreal_0(sK3(sK2,sK1),sK2) [avatar definition]
298. r1_xxreal_0(sK3(sK2,sK1),sK2) <- (19) [avatar component clause 296] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:17,thDist:-17}
299. 19 | ~3 | ~5 | ~12 | ~16 | ~17 [avatar split clause 294,199,195,94,60,50,296]
326. r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (19) [resolution 298,33] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
331. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (~1, 19) [subsumption resolution 326,42] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
332. ~v1_xxreal_0(sK2) <- (~1, 4, 19) [subsumption resolution 331,57] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
333. $false <- (~1, 3, 4, 19) [subsumption resolution 332,52] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
334. 1 | ~3 | ~4 | ~19 [avatar contradiction clause 333]
339. $false [avatar sat refutation 43,48,53,58,63,169,177,181,202,215,299,334] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:270,thDist:-270}
% SZS output end Proof for xreal_1__t229_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 6
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 239
% Active clauses: 38
% Passive clauses: 68
% Final active clauses: 33
% Final passive clauses: 27
% Main loop iterations started: 39
% 
% >>> Simplifying Inferences
% Duplicate literals: 52
% Fw subsumption resolutions: 78
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 28
% 
% >>> Generating Inferences
% Binary resolution: 83
% Unit resulting resolution: 3
% Factoring: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 6
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

