i: [1,86,0,3,1,1,0,0] t16_setfam_1
a: [3,123,0,3,1,86]
s: 123
i: [1,85,0,3,1,1,0,0] t16_setfam_1
a: [3,128,0,3,1,85]
s: 128
i: [1,87,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,133,0,2,1,87]
s: 133
i: [1,88,0,2,1,0,0,1] fc1_xboole_0
a: [3,138,0,2,1,88]
s: 138
i: [1,118,0,2,1,0,0,2] rc2_xboole_0
a: [3,148,0,2,1,118]
s: 148
i: [1,119,0,2,1,0,0,2] rc1_xboole_0
a: [3,153,0,2,1,119]
s: 153
i: [1,90,0,3,1,0,0,3] fc1_subset_1
s: 90
i: [1,112,0,3,1,0,0,3] rc2_subset_1
s: 112
i: [1,115,0,3,1,0,0,2] reflexivity_r1_tarski
s: 115
i: [1,116,0,3,1,0,0,1] reflexivity_r1_setfam_1
s: 116
i: [1,89,0,3,1,0,0,255] d2_xboole_0
a: [3,143,0,3,1,89]
s: 143
r: 133
i: [1,91,0,4,1,0,0,255] dt_o_1_0_setfam_1
s: 91
i: [1,110,0,4,1,0,0,3] existence_m1_subset_1
s: 110
i: [1,114,0,4,1,0,0,255] rc3_subset_1
s: 114
i: [1,95,0,5,2,0,0,3] rc1_subset_1
s: 95
i: [1,98,0,5,2,0,0,1] t6_boole
s: 98
d: [2,158,1,5,2,1,44,128,98]
s: 158
i: [1,107,0,5,2,0,0,2] t7_boole
s: 107
i: [1,111,0,5,1,0,0,3] rc2_subset_1
s: 111
i: [1,113,0,5,1,0,0,255] rc3_subset_1
s: 113
i: [1,97,0,6,2,0,0,255] rc4_subset_1
s: 97
i: [1,101,0,6,2,0,0,2] t1_subset
s: 101
i: [1,117,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 117
i: [1,94,0,7,2,0,0,3] rc1_subset_1
s: 94
i: [1,96,0,7,2,0,0,255] rc4_subset_1
s: 96
i: [1,105,0,7,2,0,0,2] t3_subset
s: 105
i: [1,106,0,7,3,0,0,1] t8_boole
s: 106
i: [1,99,0,8,3,0,0,3] cc1_subset_1
s: 99
d: [2,162,1,5,2,0,40,99,113]
s: 162
i: [1,102,0,8,3,0,0,2] t2_subset
s: 102
i: [1,100,0,9,3,0,0,255] cc4_subset_1
s: 100
d: [2,170,1,8,2,0,40,102,111]
d: [2,175,1,5,1,0,21,170,90]
s: 175
i: [1,109,0,9,3,0,0,2] t5_subset
s: 109
d: [2,171,1,8,2,0,40,102,113]
d: [2,176,1,5,1,0,21,171,90]
s: 176
i: [1,108,0,10,3,0,0,2] t4_subset
s: 108
i: [1,92,0,11,4,0,0,255] cc3_subset_1
s: 92
d: [2,208,1,9,3,0,40,92,113]
d: [2,215,1,5,2,0,21,208,114]
s: 215
i: [1,93,0,11,4,0,0,255] cc2_subset_1
s: 93
d: [2,185,2,5,1,0,40,175,117]
s: 185
i: [1,103,0,11,3,0,0,1] d2_setfam_1
s: 103
i: [1,104,0,11,3,0,0,1] d2_setfam_1
s: 104
d: [2,195,2,5,1,0,40,176,117]
s: 195
d: [2,164,1,6,2,0,40,99,91]
s: 164
d: [2,165,1,6,2,0,40,99,110]
s: 165
d: [2,166,1,6,2,0,40,102,91]
s: 166
d: [2,167,1,6,2,0,40,102,110]
s: 167
d: [2,180,1,6,2,0,40,100,111]
s: 180
d: [2,190,1,6,2,0,40,109,111]
s: 190
d: [2,191,1,6,2,0,40,109,113]
s: 191
d: [2,207,1,9,3,0,40,92,111]
d: [2,214,1,6,2,0,21,207,112]
s: 214
d: [2,156,1,7,3,0,44,98,98]
s: 156
d: [2,220,2,6,2,0,40,166,117]
s: 220
d: [2,163,1,7,3,0,40,99,105]
s: 163
d: [2,223,2,6,2,0,40,167,117]
s: 223
d: [2,168,1,10,3,0,40,102,94]
d: [2,173,1,7,2,0,21,168,90]
s: 173
d: [2,169,1,10,3,0,40,102,96]
d: [2,174,1,7,2,0,21,169,90]
s: 174
d: [2,172,1,10,3,0,40,102,105]
d: [2,177,1,7,2,0,21,172,90]
s: 177
d: [2,183,1,7,2,0,40,100,91]
s: 183
d: [2,184,1,7,2,0,40,100,110]
s: 184
d: [2,193,1,7,2,0,40,109,91]
s: 193
d: [2,194,1,7,2,0,40,109,110]
s: 194
d: [2,200,1,7,2,0,40,108,111]
s: 200
d: [2,201,1,7,2,0,40,108,113]
s: 201
d: [2,235,2,7,2,0,40,173,117]
s: 235
d: [2,182,1,8,3,0,40,100,105]
s: 182
d: [2,192,1,8,3,0,40,109,105]
s: 192
d: [2,256,2,5,2,0,40,182,115]
s: 256
d: [2,203,1,8,2,0,40,108,91]
s: 203
d: [2,238,2,7,2,0,40,174,117]
s: 238
d: [2,204,1,8,2,0,40,108,110]
s: 204
d: [2,219,1,8,3,0,40,103,107]
s: 219
d: [2,268,2,6,2,0,40,219,175]
s: 268
d: [2,198,1,9,3,0,40,108,94]
s: 198
d: [2,269,2,6,2,0,40,219,176]
s: 269
d: [2,199,1,9,3,0,40,108,96]
s: 199
d: [2,202,1,9,3,0,40,108,105]
s: 202
d: [2,241,2,7,2,0,40,177,117]
s: 241
d: [2,209,1,10,4,0,40,92,105]
s: 209
d: [2,272,2,7,3,0,40,219,166]
s: 272
d: [2,210,1,11,3,0,40,92,91]
s: 210
d: [2,211,1,11,3,0,40,92,110]
s: 211
d: [2,288,3,4,2,1,40,272,128]
d: [2,292,3,2,1,2,21,288,138]
a: [3,296,3,2,1,292]
s: 296
d: [2,304,4,0,0,2,52,123,296,98]
e: 304
d: [2,305,4,0,0,3,52,138,123,296,106]
e: 305
d: [2,306,4,0,0,3,52,138,123,296,106]
e: 306
f: 304,292,88,85,86
% Refutation found. Thanks to Tanya!
% SZS status Theorem for setfam_1__t16_setfam_1
% SZS output start Proof for setfam_1__t16_setfam_1
1. ! [X0] : (r1_setfam_1(X0,k1_xboole_0) => k1_xboole_0 = X0) [input t16_setfam_1]
2. ~! [X0] : (r1_setfam_1(X0,k1_xboole_0) => k1_xboole_0 = X0) [negated conjecture 1]
8. ! [X0] : ! [X1] : (r1_setfam_1(X0,X1) <=> ! [X2] : ~(! [X3] : ~(r1_tarski(X2,X3) & r2_hidden(X3,X1)) & r2_hidden(X2,X0))) [input d2_setfam_1]
14. ! [X0] : m1_subset_1(o_1_0_setfam_1(X0),X0) [input dt_o_1_0_setfam_1]
17. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
27. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
31. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
32. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
35. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 27]
36. ! [X0,X1] : (r1_setfam_1(X0,X1) <=> ! [X2] : ~(! [X3] : ~(r1_tarski(X2,X3) & r2_hidden(X3,X1)) & r2_hidden(X2,X0))) [flattening 8]
39. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 32]
44. ! [X0,X1] : (r1_setfam_1(X0,X1) => ! [X2] : ~(! [X3] : ~(r1_tarski(X2,X3) & r2_hidden(X3,X1)) & r2_hidden(X2,X0))) [unused predicate definition removal 36]
46. ? [X0] : (k1_xboole_0 != X0 & r1_setfam_1(X0,k1_xboole_0)) [ennf transformation 2]
53. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 31]
57. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 35]
58. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 57]
59. ! [X0,X1] : (! [X2] : (? [X3] : (r1_tarski(X2,X3) & r2_hidden(X3,X1)) | ~r2_hidden(X2,X0)) | ~r1_setfam_1(X0,X1)) [ennf transformation 44]
62. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 39]
67. ? [X0] : (k1_xboole_0 != X0 & r1_setfam_1(X0,k1_xboole_0)) => (k1_xboole_0 != sK0 & r1_setfam_1(sK0,k1_xboole_0)) [choice axiom]
68. k1_xboole_0 != sK0 & r1_setfam_1(sK0,k1_xboole_0) [skolemisation 46,67]
73. ! [X2,X1] : (? [X3] : (r1_tarski(X2,X3) & r2_hidden(X3,X1)) => (r1_tarski(X2,sK3(X1,X2)) & r2_hidden(sK3(X1,X2),X1))) [choice axiom]
74. ! [X0,X1] : (! [X2] : ((r1_tarski(X2,sK3(X1,X2)) & r2_hidden(sK3(X1,X2),X1)) | ~r2_hidden(X2,X0)) | ~r1_setfam_1(X0,X1)) [skolemisation 59,73]
85. r1_setfam_1(sK0,k1_xboole_0) [cnf transformation 68] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
86. k1_xboole_0 != sK0 [cnf transformation 68] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
88. v1_xboole_0(k1_xboole_0) [cnf transformation 17] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
91. m1_subset_1(o_1_0_setfam_1(X0),X0) [cnf transformation 14] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
98. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 53] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
102. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 58] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
103. r2_hidden(sK3(X1,X2),X1) | ~r2_hidden(X2,X0) | ~r1_setfam_1(X0,X1) [cnf transformation 74] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
107. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 62] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
121. 1 <=> k1_xboole_0 = sK0 [avatar definition]
123. k1_xboole_0 != sK0 <- (~1) [avatar component clause 121] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
124. ~1 [avatar split clause 86,121]
126. 2 <=> r1_setfam_1(sK0,k1_xboole_0) [avatar definition]
128. r1_setfam_1(sK0,k1_xboole_0) <- (2) [avatar component clause 126] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
129. 2 [avatar split clause 85,126]
136. 4 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
138. v1_xboole_0(k1_xboole_0) <- (4) [avatar component clause 136] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
139. 4 [avatar split clause 88,136]
166. r2_hidden(o_1_0_setfam_1(X0),X0) | v1_xboole_0(X0) [resolution 102,91] {a:1,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
219. ~r2_hidden(X6,X7) | ~r1_setfam_1(X7,X8) | ~v1_xboole_0(X8) [resolution 103,107] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
272. ~r1_setfam_1(X15,X16) | ~v1_xboole_0(X16) | v1_xboole_0(X15) [resolution 219,166] {a:2,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
288. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK0) <- (2) [resolution 272,128] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
292. v1_xboole_0(sK0) <- (2, 4) [subsumption resolution 288,138] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
294. 8 <=> v1_xboole_0(sK0) [avatar definition]
296. v1_xboole_0(sK0) <- (8) [avatar component clause 294] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
297. 8 | ~2 | ~4 [avatar split clause 292,136,126,294]
304. $false <- (~1, 8) [unit resulting resolution 123,296,98] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
307. 1 | ~8 [avatar contradiction clause 304]
310. $false [avatar sat refutation 124,129,139,297,307] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:71,thDist:-71}
% SZS output end Proof for setfam_1__t16_setfam_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 32
% >>> Preprocessing
% Introduced skolems: 9
% Unused predicate definitions: 2
% 
% >>> Saturation
% Initial clauses: 35
% Generated clauses: 186
% Active clauses: 79
% Passive clauses: 115
% Final active clauses: 78
% Final passive clauses: 33
% Main loop iterations started: 82
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 10
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 16
% Forward subsumptions: 33
% 
% >>> Generating Inferences
% Binary resolution: 123
% Unit resulting resolution: 3
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 7
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% passive container maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

