i: [1,283,0,2,1,1,0,0] t222_xreal_1
a: [3,449,0,2,1,283]
s: 449
i: [1,284,0,3,1,1,0,0] t222_xreal_1
a: [3,444,0,3,1,284]
s: 444
i: [1,285,0,5,1,1,0,0] t222_xreal_1
a: [3,439,0,5,1,285]
s: 439
i: [1,286,0,2,1,0,0,2] fc1_numbers
a: [3,454,0,2,1,286]
s: 454
i: [1,287,0,2,1,0,0,1] spc3_boole
a: [3,459,0,2,1,287]
s: 459
i: [1,288,0,2,1,0,0,2] spc1_boole
a: [3,464,0,2,1,288]
s: 464
i: [1,289,0,2,1,0,0,255] fc1_xboole_0
a: [3,469,0,2,1,289]
s: 469
i: [1,316,0,2,1,0,0,255] fc6_ordinal1
a: [3,594,0,2,1,316]
s: 594
i: [1,315,0,2,1,0,0,255] fc6_ordinal1
a: [3,599,0,2,1,315]
s: 599
i: [1,317,0,2,1,0,0,2] spc3_numerals
a: [3,619,0,2,1,317]
s: 619
i: [1,321,0,2,1,0,0,2] spc1_numerals
a: [3,639,0,2,1,321]
s: 639
i: [1,403,0,2,1,0,0,2] rc2_xboole_0
a: [3,651,0,2,1,403]
s: 651
i: [1,404,0,2,1,0,0,255] rc4_ordinal1
a: [3,656,0,2,1,404]
s: 656
i: [1,405,0,2,1,0,0,255] rc1_ordinal1
a: [3,661,0,2,1,405]
s: 661
i: [1,406,0,2,1,0,0,1] rc1_xreal_0
a: [3,666,0,2,1,406]
s: 666
i: [1,407,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,671,0,2,1,407]
s: 671
i: [1,408,0,2,1,0,0,255] rc1_xxreal_0
a: [3,676,0,2,1,408]
s: 676
i: [1,409,0,2,1,0,0,2] rc1_xboole_0
a: [3,681,0,2,1,409]
s: 681
i: [1,413,0,2,1,0,0,255] rc2_ordinal1
a: [3,686,0,2,1,413]
s: 686
i: [1,412,0,2,1,0,0,255] rc2_ordinal1
a: [3,691,0,2,1,412]
s: 691
i: [1,411,0,2,1,0,0,255] rc2_ordinal1
a: [3,696,0,2,1,411]
s: 696
i: [1,410,0,2,1,0,0,255] rc2_ordinal1
a: [3,701,0,2,1,410]
s: 701
i: [1,415,0,2,1,0,0,255] rc5_ordinal1
a: [3,706,0,2,1,415]
s: 706
i: [1,414,0,2,1,0,0,255] rc5_ordinal1
a: [3,711,0,2,1,414]
s: 711
i: [1,417,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,716,0,2,1,417]
s: 716
i: [1,416,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,721,0,2,1,416]
s: 721
i: [1,421,0,2,1,0,0,255] rc3_xreal_0
a: [3,726,0,2,1,421]
s: 726
i: [1,420,0,2,1,0,0,255] rc3_xreal_0
a: [3,731,0,2,1,420]
s: 731
i: [1,419,0,2,1,0,0,255] rc3_xreal_0
a: [3,736,0,2,1,419]
s: 736
i: [1,418,0,2,1,0,0,255] rc3_xreal_0
a: [3,741,0,2,1,418]
s: 741
i: [1,425,0,2,1,0,0,255] rc2_xreal_0
a: [3,746,0,2,1,425]
s: 746
i: [1,424,0,2,1,0,0,255] rc2_xreal_0
a: [3,751,0,2,1,424]
s: 751
i: [1,423,0,2,1,0,0,255] rc2_xreal_0
a: [3,756,0,2,1,423]
s: 756
i: [1,422,0,2,1,0,0,255] rc2_xreal_0
a: [3,761,0,2,1,422]
s: 761
i: [1,427,0,2,1,0,0,255] rc3_xxreal_0
a: [3,766,0,2,1,427]
s: 766
i: [1,426,0,2,1,0,0,255] rc3_xxreal_0
a: [3,771,0,2,1,426]
s: 771
i: [1,429,0,2,1,0,0,255] rc2_xxreal_0
a: [3,776,0,2,1,429]
s: 776
i: [1,428,0,2,1,0,0,255] rc2_xxreal_0
a: [3,781,0,2,1,428]
s: 781
i: [1,433,0,2,1,0,0,255] rc4_xreal_0
a: [3,786,0,2,1,433]
s: 786
i: [1,432,0,2,1,0,0,255] rc4_xreal_0
a: [3,791,0,2,1,432]
s: 791
i: [1,431,0,2,1,0,0,255] rc4_xreal_0
a: [3,796,0,2,1,431]
s: 796
i: [1,430,0,2,1,0,0,255] rc4_xreal_0
a: [3,801,0,2,1,430]
s: 801
i: [1,435,0,2,1,0,0,255] rc4_xxreal_0
a: [3,806,0,2,1,435]
s: 806
i: [1,434,0,2,1,0,0,255] rc4_xxreal_0
a: [3,811,0,2,1,434]
s: 811
i: [1,378,0,3,1,0,0,255] reflexivity_r1_tarski
s: 378
i: [1,290,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r3_r1
a: [3,474,0,3,1,290]
s: 474
i: [1,291,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,479,0,3,1,291]
s: 479
i: [1,292,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,484,0,3,1,292]
s: 484
r: 469
d: [2,813,0,2,1,2,23,469,484]
a: [3,817,0,2,1,813]
s: 817
i: [1,293,0,3,1,0,0,255] t1_numerals
a: [3,489,0,3,1,293]
s: 489
i: [1,294,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r3_r3
a: [3,494,0,3,1,294]
s: 494
i: [1,295,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r3
a: [3,499,0,3,1,295]
s: 499
i: [1,296,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,504,0,3,1,296]
s: 504
i: [1,320,0,3,1,0,0,2] spc3_numerals
a: [3,604,0,3,1,320]
s: 604
i: [1,319,0,3,1,0,0,2] spc3_numerals
a: [3,609,0,3,1,319]
s: 609
i: [1,324,0,3,1,0,0,2] spc1_numerals
a: [3,624,0,3,1,324]
s: 624
i: [1,323,0,3,1,0,0,2] spc1_numerals
a: [3,629,0,3,1,323]
s: 629
d: [2,819,0,3,1,2,22,489,484]
a: [3,823,0,3,1,819]
s: 823
d: [2,825,0,3,1,2,22,609,479]
a: [3,829,0,3,1,825]
s: 829
d: [2,831,0,3,1,2,22,629,479]
a: [3,835,0,3,1,831]
s: 835
i: [1,326,0,4,2,0,0,255] cc6_ordinal1
s: 326
d: [2,837,1,2,1,1,40,326,656]
a: [3,842,1,2,1,837]
s: 842
i: [1,327,0,4,2,0,0,255] cc2_xreal_0
s: 327
d: [2,838,1,2,1,1,40,326,706]
a: [3,847,1,2,1,838]
s: 847
i: [1,328,0,4,2,0,0,255] cc2_xcmplx_0
s: 328
i: [1,329,0,4,2,0,0,255] cc2_xxreal_0
s: 329
d: [2,849,1,2,1,1,40,327,656]
a: [3,854,1,2,1,849]
s: 854
i: [1,330,0,4,2,0,0,255] cc1_ordinal1
s: 330
d: [2,850,1,2,1,1,40,327,706]
a: [3,859,1,2,1,850]
s: 859
i: [1,331,0,4,2,0,0,255] cc1_ordinal1
s: 331
i: [1,333,0,4,2,0,0,2] cc3_xreal_0
s: 333
d: [2,935,1,2,1,1,40,333,449]
a: [3,945,1,2,1,935]
s: 945
i: [1,334,0,4,2,0,0,255] cc4_xreal_0
s: 334
d: [2,954,1,2,1,1,40,334,449]
a: [3,964,1,2,1,954]
s: 964
i: [1,347,0,4,2,0,0,255] cc9_ordinal1
s: 347
i: [1,348,0,4,2,0,0,255] cc7_ordinal1
s: 348
d: [2,861,1,2,1,1,40,328,656]
a: [3,866,1,2,1,861]
s: 866
i: [1,349,0,4,2,0,0,255] cc3_ordinal1
s: 349
d: [2,862,1,2,1,1,40,328,706]
a: [3,871,1,2,1,862]
s: 871
i: [1,377,0,4,1,0,0,3] existence_m1_subset_1
s: 377
i: [1,297,0,4,1,0,0,2] dt_k5_numbers
a: [3,509,0,4,1,297]
s: 509
d: [2,873,1,2,1,1,40,329,656]
a: [3,878,1,2,1,873]
s: 878
i: [1,298,0,4,1,0,0,1] dt_k6_numbers
a: [3,514,0,4,1,298]
s: 514
d: [2,874,1,2,1,1,40,329,706]
a: [3,883,1,2,1,874]
s: 883
i: [1,318,0,4,1,0,0,2] spc3_numerals
a: [3,614,0,4,1,318]
s: 614
i: [1,322,0,4,1,0,0,2] spc1_numerals
a: [3,634,0,4,1,322]
s: 634
d: [2,885,1,2,1,1,40,330,594]
a: [3,893,1,2,1,885]
s: 893
d: [2,1045,0,4,1,2,22,509,479]
a: [3,1049,0,4,1,1045]
s: 1049
d: [2,887,1,2,1,1,40,330,661]
a: [3,903,1,2,1,887]
s: 903
d: [2,1051,0,4,1,2,22,514,479]
a: [3,1055,0,4,1,1051]
s: 1055
d: [2,1057,0,4,1,2,22,614,479]
a: [3,1061,0,4,1,1057]
s: 1061
d: [2,910,1,2,1,1,40,331,594]
a: [3,918,1,2,1,910]
s: 918
d: [2,1063,0,4,1,2,22,634,479]
a: [3,1067,0,4,1,1063]
s: 1067
d: [2,912,1,2,1,1,40,331,661]
a: [3,928,1,2,1,912]
s: 928
i: [1,350,0,5,2,0,0,255] t6_boole
s: 350
i: [1,351,0,5,2,0,0,255] cc8_ordinal1
s: 351
d: [2,937,1,2,1,1,40,333,666]
a: [3,951,1,2,1,937]
s: 951
i: [1,352,0,5,2,0,0,2] cc1_xreal_0
s: 352
d: [2,956,1,2,1,1,40,334,666]
a: [3,970,1,2,1,956]
s: 970
i: [1,353,0,5,2,0,0,2] cc1_xcmplx_0
s: 353
i: [1,374,0,5,2,0,0,255] t7_boole
s: 374
d: [2,973,1,2,1,1,40,347,817]
a: [3,980,1,2,1,973]
s: 980
i: [1,299,0,5,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r3_rn1d3
a: [3,519,0,5,1,299]
s: 519
d: [2,974,1,2,1,1,40,347,681]
a: [3,985,1,2,1,974]
s: 985
i: [1,300,0,5,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_rn1d3
a: [3,524,0,5,1,300]
s: 524
i: [1,301,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r3_r1_r3
a: [3,529,0,5,1,301]
s: 529
d: [2,975,1,2,1,1,40,347,801]
a: [3,990,1,2,1,975]
s: 990
i: [1,302,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_r3_r3
a: [3,534,0,5,1,302]
s: 534
d: [2,976,1,2,1,1,40,347,811]
a: [3,995,1,2,1,976]
s: 995
i: [1,303,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,539,0,5,1,303]
s: 539
i: [1,304,0,5,1,0,0,1] rqRealDiv__k7_xcmplx_0__r3_r3_r1
a: [3,544,0,5,1,304]
s: 544
d: [2,997,1,2,1,1,40,348,817]
a: [3,1004,1,2,1,997]
s: 1004
i: [1,305,0,5,1,0,0,2] rqRealDiv__k7_xcmplx_0__r1_r1_r1
a: [3,549,0,5,1,305]
s: 549
d: [2,998,1,2,1,1,40,348,681]
a: [3,1009,1,2,1,998]
s: 1009
i: [1,306,0,5,1,0,0,1] rqLessOrEqual__r1_xxreal_0__rn1d3_r3
a: [3,554,0,5,1,306]
s: 554
i: [1,307,0,5,1,0,0,1] rqLessOrEqual__r1_xxreal_0__rn1d3_r1
a: [3,559,0,5,1,307]
s: 559
d: [2,999,1,2,1,1,40,348,801]
a: [3,1014,1,2,1,999]
s: 1014
i: [1,399,0,7,3,0,0,255] reflexivity_r1_xxreal_0
a: [3,646,0,5,2,399]
s: 646
d: [2,1000,1,2,1,1,40,348,811]
a: [3,1019,1,2,1,1000]
s: 1019
d: [2,1069,0,5,2,1,22,350,484]
s: 1069
i: [1,358,0,6,3,0,0,255] cc2_ordinal1
s: 358
d: [2,1188,1,5,2,2,44,444,1069]
s: 1188
i: [1,361,0,6,3,0,0,255] cc5_xxreal_0
s: 361
d: [2,1197,2,4,2,3,40,1188,646]
d: [2,1198,2,2,1,4,21,1197,964]
a: [3,1202,2,2,1,1198]
s: 1202
i: [1,363,0,6,3,0,0,255] cc5_xxreal_0
s: 363
i: [1,364,0,6,3,0,0,255] cc3_xxreal_0
s: 364
d: [2,1187,1,7,2,2,44,439,1069]
s: 1187
i: [1,366,0,6,3,0,0,255] cc3_xxreal_0
s: 366
d: [2,1021,1,2,1,1,40,349,817]
a: [3,1028,1,2,1,1021]
s: 1028
i: [1,368,0,6,3,0,0,255] cc7_xxreal_0
s: 368
i: [1,369,0,6,3,0,0,255] cc7_xxreal_0
s: 369
d: [2,1022,1,2,1,1,40,349,681]
a: [3,1033,1,2,1,1022]
s: 1033
i: [1,370,0,6,2,0,0,255] t1_subset
s: 370
d: [2,1023,1,2,1,1,40,349,801]
a: [3,1038,1,2,1,1023]
s: 1038
i: [1,379,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 379
i: [1,332,0,7,3,0,0,255] cc5_ordinal1
s: 332
d: [2,1024,1,2,1,1,40,349,811]
a: [3,1043,1,2,1,1024]
s: 1043
i: [1,343,0,7,2,0,0,1] t2_arithm
s: 343
d: [2,1071,1,2,1,1,40,351,829]
a: [3,1078,1,2,1,1071]
s: 1078
i: [1,344,0,7,2,0,0,1] t5_arithm
s: 344
i: [1,345,0,7,2,0,0,2] t3_arithm
s: 345
d: [2,1072,1,2,1,1,40,351,835]
a: [3,1083,1,2,1,1072]
s: 1083
i: [1,346,0,7,2,0,0,2] t6_arithm
s: 346
d: [2,1090,1,2,1,1,40,352,604]
a: [3,1096,1,2,1,1090]
s: 1096
i: [1,372,0,7,2,0,0,255] t3_subset
s: 372
i: [1,373,0,7,3,0,0,2] t8_boole
s: 373
d: [2,1091,1,2,1,1,40,352,624]
a: [3,1101,1,2,1,1091]
s: 1101
i: [1,308,0,7,1,0,0,1] rqRealDiv__k7_xcmplx_0__r1_rn1d3_r3
a: [3,564,0,7,1,308]
s: 564
d: [2,1108,1,2,1,1,40,353,604]
a: [3,1114,1,2,1,1108]
s: 1114
i: [1,309,0,7,1,0,0,2] rqRealMult__k3_xcmplx_0__r3_rn1d3_r1
a: [3,569,0,7,1,309]
s: 569
i: [1,310,0,7,1,0,0,2] rqRealMult__k3_xcmplx_0__rn1d3_r3_r1
a: [3,574,0,7,1,310]
s: 574
d: [2,1109,1,2,1,1,40,353,624]
a: [3,1119,1,2,1,1109]
s: 1119
i: [1,312,0,7,1,0,0,1] rqLessOrEqual__r1_xxreal_0__rn1d3_rn1d3
a: [3,579,0,7,1,312]
s: 579
d: [2,1204,1,4,2,1,40,361,731]
d: [2,1206,1,2,1,2,21,1204,736]
a: [3,1210,1,2,1,1206]
s: 1210
i: [1,325,0,8,3,0,0,255] cc10_ordinal1
s: 325
i: [1,355,0,8,4,0,0,255] cc4_xxreal_0
s: 355
d: [2,1381,2,6,3,1,40,355,964]
d: [2,1392,2,4,2,2,21,1381,1202]
a: [3,1396,2,2,1,1392]
s: 1396
i: [1,357,0,8,4,0,0,255] cc6_xxreal_0
s: 357
a: [3,1399,2,2,1,1392]
s: 1399
i: [1,359,0,8,4,0,0,255] cc8_xxreal_0
s: 359
i: [1,371,0,8,3,0,0,255] t2_subset
s: 371
d: [2,1205,1,4,2,1,40,361,766]
d: [2,1212,1,2,1,2,21,1205,771]
a: [3,1216,1,2,1,1212]
s: 1216
i: [1,394,0,8,3,0,0,2] fc6_xreal_0
s: 394
d: [2,1218,1,4,2,1,40,363,619]
a: [3,1224,1,2,1,1218]
s: 1224
i: [1,395,0,8,3,0,0,1] fc8_xreal_0
s: 395
i: [1,396,0,8,3,0,0,2] fc3_xcmplx_0
s: 396
a: [3,1229,1,2,1,1218]
s: 1229
i: [1,397,0,8,3,0,0,2] fc5_xcmplx_0
s: 397
d: [2,1219,1,4,2,1,40,363,639]
a: [3,1233,1,2,1,1219]
s: 1233
i: [1,376,0,9,3,0,0,255] t5_subset
s: 376
i: [1,313,0,9,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_rn1d3_rn1d3
a: [3,584,0,9,1,313]
s: 584
a: [3,1238,1,2,1,1219]
s: 1238
i: [1,314,0,9,1,0,0,2] rqRealMult__k3_xcmplx_0__rn1d3_r1_rn1d3
a: [3,589,0,9,1,314]
s: 589
d: [2,1220,1,4,2,1,40,363,751]
d: [2,1240,1,2,1,2,21,1220,756]
a: [3,1244,1,2,1,1240]
s: 1244
i: [1,375,0,10,3,0,0,255] t4_subset
s: 375
i: [1,400,0,10,4,0,0,255] connectedness_r1_xxreal_0
s: 400
d: [2,1552,1,7,3,1,40,400,444]
d: [2,1559,1,5,2,2,21,1552,964]
a: [3,1563,1,3,1,1559]
s: 1563
i: [1,337,0,11,5,0,0,1] t2_real
s: 337
d: [2,1571,2,14,5,2,40,337,1187]
a: [3,1587,2,4,1,1571]
s: 1587
d: [2,1601,3,0,0,3,52,449,1096,1587,395]
e: 1601
d: [2,1602,3,4,2,1,40,1587,395]
d: [2,1604,3,2,1,2,21,1602,1096]
d: [2,1605,3,0,0,3,21,1604,449]
e: 1605
r: 1587
i: [1,338,0,11,5,0,0,1] t1_real
s: 338
i: [1,339,0,11,5,0,0,1] t8_real
s: 339
a: [3,1586,2,4,1,1571]
s: 1586
i: [1,340,0,11,5,0,0,1] t7_real
s: 340
a: [3,1591,2,4,1,1571]
s: 1591
i: [1,341,0,11,5,0,0,1] t3_real
s: 341
i: [1,342,0,11,5,0,0,1] t4_real
s: 342
d: [2,1611,2,14,5,2,40,338,1187]
a: [3,1630,2,4,1,1611]
s: 1630
i: [1,398,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 398
d: [2,1646,3,4,1,1,40,1586,334]
a: [3,1651,3,4,1,1646]
s: 1651
i: [1,380,0,12,5,0,0,2] fc26_xreal_0
s: 380
i: [1,381,0,12,5,0,0,2] fc33_xreal_0
s: 381
d: [2,1722,3,0,0,5,52,1399,449,1229,1096,1591,381]
e: 1722
d: [2,1723,3,8,4,1,40,381,1591]
d: [2,1732,3,6,3,2,21,1723,1096]
d: [2,1733,3,4,2,3,21,1732,1229]
d: [2,1734,3,2,1,4,21,1733,449]
d: [2,1735,3,0,0,5,21,1734,1399]
e: 1735
r: 1396
r: 1399
r: 1591
r: 1630
d: [2,1572,1,12,4,1,40,337,439]
a: [3,1599,1,2,1,1572]
s: 1599
i: [1,382,0,12,5,0,0,2] fc32_xreal_0
s: 382
d: [2,1612,1,12,4,1,40,338,439]
a: [3,1635,1,2,1,1612]
s: 1635
r: 1635
i: [1,383,0,12,5,0,0,2] fc31_xreal_0
s: 383
i: [1,388,0,12,5,0,0,2] fc8_xcmplx_0
s: 388
a: [3,1634,1,2,1,1612]
s: 1634
i: [1,389,0,12,5,0,0,2] fc9_xcmplx_0
s: 389
s: 1396
i: [1,390,0,12,5,0,0,2] fc23_xreal_0
s: 390
i: [1,391,0,12,5,0,0,2] fc24_xreal_0
s: 391
s: 1399
i: [1,392,0,12,5,0,0,2] fc25_xreal_0
s: 392
a: [3,1590,2,4,1,1571]
s: 1590
i: [1,393,0,12,5,0,0,2] fc34_xreal_0
s: 393
i: [1,335,0,13,6,0,0,1] t5_real
s: 335
s: 1630
i: [1,336,0,13,6,0,0,1] t6_real
s: 336
d: [2,1752,2,4,2,2,44,1599,1069]
s: 1752
r: 361
i: [1,387,0,14,4,0,0,2] existence_m2_subset_1
s: 387
i: [1,384,0,15,5,0,0,2] dt_m2_subset_1
s: 384
d: [2,1806,2,4,2,2,44,1634,1069]
s: 1806
r: 364
i: [1,385,0,15,5,0,0,2] redefinition_m2_subset_1
s: 385
d: [2,1647,3,4,1,1,40,1586,333]
a: [3,1656,3,4,1,1647]
s: 1656
i: [1,386,0,15,5,0,0,2] redefinition_m2_subset_1
s: 386
i: [1,401,0,17,4,0,0,2] spc4_arithm
s: 401
d: [2,1658,3,8,2,1,40,1591,361]
a: [3,1661,3,4,1,1658]
s: 1661
d: [2,2043,4,0,0,5,52,1202,945,459,1114,1661,389]
e: 2043
d: [2,2044,4,8,4,1,40,1661,389]
d: [2,2049,4,6,3,2,21,2044,1114]
d: [2,2050,4,4,2,3,21,2049,459]
d: [2,2051,4,2,1,4,21,2050,945]
d: [2,2052,4,0,0,5,21,2051,1202]
e: 2052
r: 1396
r: 1399
r: 1590
r: 1752
r: 1599
r: 1661
s: 361
a: [3,1598,1,2,1,1572]
s: 1598
r: 646
r: 1202
r: 1598
d: [2,1070,1,2,1,1,40,351,823]
d: [2,1126,2,2,1,1,40,1004,329]
d: [2,1128,2,2,1,1,40,1004,327]
d: [2,1360,2,2,1,1,40,1096,334]
d: [2,1642,1,8,4,1,40,339,444]
d: [2,1742,1,6,3,2,21,1642,449]
d: [2,1760,2,4,2,1,40,1635,364]
d: [2,1763,2,2,1,2,21,1760,817]
d: [2,1817,3,4,2,1,40,1396,363]
d: [2,1818,3,2,1,2,21,1817,964]
d: [2,1914,1,16,5,1,40,336,439]
d: [2,1929,1,14,4,2,21,1914,1634]
d: [2,2057,1,10,3,3,21,1929,1586]
d: [2,2073,2,4,2,1,40,1598,361]
d: [2,2075,2,2,1,2,21,2073,817]
d: [2,1557,1,7,3,0,42,400]
d: [2,1558,1,5,2,0,18,1557]
f: 1558,2075,2057,2052,1818,1763,1734,1742,1605,1360,1218,1198,1128,1126,1108,1090,1070,954,935,819,813,317,320,293,292,289,287,283,284,285
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t222_xreal_1
% SZS output start Proof for xreal_1__t222_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ~(r1_xxreal_0(k7_xcmplx_0(X0,3),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers))) [input t222_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ~(r1_xxreal_0(k7_xcmplx_0(X0,3),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers))) [negated conjecture 1]
12. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
13. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
16. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
18. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
25. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc33_xreal_0]
32. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
40. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0))))) [input t6_real]
42. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
52. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xcmplx_0(X0)) [input cc1_xcmplx_0]
53. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xreal_0(X0)) [input cc1_xreal_0]
55. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
62. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
69. ! [X0,X1] : ((v1_xcmplx_0(X1) & ~v1_xboole_0(X1) & v1_xcmplx_0(X0) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k7_xcmplx_0(X0,X1))) [input fc9_xcmplx_0]
88. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
93. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
94. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
98. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
99. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
101. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
122. ~v1_xboole_0(3) [input spc3_boole]
124. m1_subset_1(3,k1_numbers) & m1_subset_1(3,k5_numbers) & m2_subset_1(3,k1_numbers,k5_numbers) & v2_xxreal_0(3) [input spc3_numerals]
138. ? [X0] : ((r1_xxreal_0(k7_xcmplx_0(X0,3),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers)) & v1_xreal_0(X0)) [ennf transformation 2]
139. ? [X0] : (r1_xxreal_0(k7_xcmplx_0(X0,3),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X0)) [flattening 138]
142. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
144. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 13]
147. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 98]
148. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 99]
151. ! [X0] : (! [X1] : (((v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0)) | ~r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 40]
152. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 151]
157. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 42]
158. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 157]
172. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 88]
173. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 18]
174. ! [X0] : (v1_xreal_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 53]
175. ! [X0] : (v1_xcmplx_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 52]
184. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 16]
185. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 184]
186. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 55]
187. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 186]
202. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 25]
203. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 202]
216. ! [X0,X1] : (~v1_xboole_0(k7_xcmplx_0(X0,X1)) | (~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0))) [ennf transformation 69]
217. ! [X0,X1] : (~v1_xboole_0(k7_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0)) [flattening 216]
228. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 101]
229. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 228]
238. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 93]
239. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 238]
244. ? [X0] : (r1_xxreal_0(k7_xcmplx_0(X0,3),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X0)) => (r1_xxreal_0(k7_xcmplx_0(sK0,3),k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK0)) [choice axiom]
245. r1_xxreal_0(k7_xcmplx_0(sK0,3),k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK0) [skolemisation 139,244]
283. v1_xreal_0(sK0) [cnf transformation 245] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
284. ~r1_xxreal_0(sK0,k6_numbers) [cnf transformation 245] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
285. r1_xxreal_0(k7_xcmplx_0(sK0,3),k6_numbers) [cnf transformation 245] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
287. ~v1_xboole_0(3) [cnf transformation 122] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
289. v1_xboole_0(k1_xboole_0) [cnf transformation 62] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
292. k6_numbers = k1_xboole_0 [cnf transformation 94] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
293. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 32] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
317. v2_xxreal_0(3) [cnf transformation 124] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
320. m1_subset_1(3,k1_numbers) [cnf transformation 124] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
327. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 142] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
329. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 144] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
333. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 147] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
334. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 148] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
336. ~r1_xxreal_0(X0,X1) | v2_xxreal_0(X1) | v1_xboole_0(X0) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 152] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
339. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 158] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
350. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 172] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
351. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 173] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
352. ~m1_subset_1(X0,k1_numbers) | v1_xreal_0(X0) [cnf transformation 174] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
353. ~m1_subset_1(X0,k1_numbers) | v1_xcmplx_0(X0) [cnf transformation 175] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
361. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 185] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
363. ~v2_xxreal_0(X0) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 185] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
364. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 187] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
381. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 203] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
389. ~v1_xboole_0(k7_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) [cnf transformation 217] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
395. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 229] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
400. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 239] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
437. 1 <=> r1_xxreal_0(k7_xcmplx_0(sK0,3),k6_numbers) [avatar definition]
439. r1_xxreal_0(k7_xcmplx_0(sK0,3),k6_numbers) <- (1) [avatar component clause 437] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
440. 1 [avatar split clause 285,437]
442. 2 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
444. ~r1_xxreal_0(sK0,k6_numbers) <- (~2) [avatar component clause 442] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
445. ~2 [avatar split clause 284,442]
447. 3 <=> v1_xreal_0(sK0) [avatar definition]
449. v1_xreal_0(sK0) <- (3) [avatar component clause 447] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
450. 3 [avatar split clause 283,447]
457. 5 <=> v1_xboole_0(3) [avatar definition]
459. ~v1_xboole_0(3) <- (~5) [avatar component clause 457] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
460. ~5 [avatar split clause 287,457]
467. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
469. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 467] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
470. 7 [avatar split clause 289,467]
482. 10 <=> k6_numbers = k1_xboole_0 [avatar definition]
484. k6_numbers = k1_xboole_0 <- (10) [avatar component clause 482] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
485. 10 [avatar split clause 292,482]
487. 11 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
489. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (11) [avatar component clause 487] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
490. 11 [avatar split clause 293,487]
602. 34 <=> m1_subset_1(3,k1_numbers) [avatar definition]
604. m1_subset_1(3,k1_numbers) <- (34) [avatar component clause 602] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
605. 34 [avatar split clause 320,602]
617. 37 <=> v2_xxreal_0(3) [avatar definition]
619. v2_xxreal_0(3) <- (37) [avatar component clause 617] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
620. 37 [avatar split clause 317,617]
645. 43 <=> ! [X0] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0)) [avatar definition]
646. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) <- (43) [avatar component clause 645] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
813. v1_xboole_0(k6_numbers) <- (7, 10) [backward demodulation 469,484] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
815. 77 <=> v1_xboole_0(k6_numbers) [avatar definition]
817. v1_xboole_0(k6_numbers) <- (77) [avatar component clause 815] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
818. 77 | ~7 | ~10 [avatar split clause 813,482,467,815]
819. m1_subset_1(k6_numbers,k4_ordinal1) <- (10, 11) [forward demodulation 489,484] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
821. 78 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
823. m1_subset_1(k6_numbers,k4_ordinal1) <- (78) [avatar component clause 821] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
824. 78 | ~10 | ~11 [avatar split clause 819,487,482,821]
935. v1_xcmplx_0(sK0) <- (3) [resolution 333,449] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
943. 97 <=> v1_xcmplx_0(sK0) [avatar definition]
945. v1_xcmplx_0(sK0) <- (97) [avatar component clause 943] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
946. 97 | ~3 [avatar split clause 935,447,943]
954. v1_xxreal_0(sK0) <- (3) [resolution 334,449] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
962. 99 <=> v1_xxreal_0(sK0) [avatar definition]
964. v1_xxreal_0(sK0) <- (99) [avatar component clause 962] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
965. 99 | ~3 [avatar split clause 954,447,962]
1002. 105 <=> v7_ordinal1(k6_numbers) [avatar definition]
1004. v7_ordinal1(k6_numbers) <- (105) [avatar component clause 1002] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1069. k6_numbers = X0 | ~v1_xboole_0(X0) <- (10) [forward demodulation 350,484] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1070. v7_ordinal1(k6_numbers) <- (78) [resolution 351,823] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1074. 105 | ~78 [avatar split clause 1070,821,1002]
1090. v1_xreal_0(3) <- (34) [resolution 352,604] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
1094. 120 <=> v1_xreal_0(3) [avatar definition]
1096. v1_xreal_0(3) <- (120) [avatar component clause 1094] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1097. 120 | ~34 [avatar split clause 1090,602,1094]
1108. v1_xcmplx_0(3) <- (34) [resolution 353,604] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
1112. 123 <=> v1_xcmplx_0(3) [avatar definition]
1114. v1_xcmplx_0(3) <- (123) [avatar component clause 1112] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1115. 123 | ~34 [avatar split clause 1108,602,1112]
1126. v1_xxreal_0(k6_numbers) <- (105) [resolution 1004,329] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1128. v1_xreal_0(k6_numbers) <- (105) [resolution 1004,327] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1131. 126 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1134. 126 | ~105 [avatar split clause 1126,1002,1131]
1141. 128 <=> v1_xreal_0(k6_numbers) [avatar definition]
1144. 128 | ~105 [avatar split clause 1128,1002,1141]
1188. ~r1_xxreal_0(sK0,X0) | ~v1_xboole_0(X0) <- (~2, 10) [superposition 444,1069] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1197. ~v1_xboole_0(sK0) | ~v1_xxreal_0(sK0) <- (~2, 10, 43) [resolution 1188,646] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1198. ~v1_xboole_0(sK0) <- (~2, 10, 43, 99) [subsumption resolution 1197,964] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1200. 134 <=> v1_xboole_0(sK0) [avatar definition]
1202. ~v1_xboole_0(sK0) <- (~134) [avatar component clause 1200] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1203. ~134 | 2 | ~10 | ~43 | ~99 [avatar split clause 1198,962,645,482,442,1200]
1218. ~v3_xxreal_0(3) | ~v1_xxreal_0(3) <- (37) [resolution 363,619] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1223. 137 <=> v1_xxreal_0(3) [avatar definition]
1227. 138 <=> v3_xxreal_0(3) [avatar definition]
1229. ~v3_xxreal_0(3) <- (~138) [avatar component clause 1227] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1230. ~137 | ~138 | ~37 [avatar split clause 1218,617,1227,1223]
1360. v1_xxreal_0(3) <- (120) [resolution 1096,334] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1362. 137 | ~120 [avatar split clause 1360,1094,1223]
1394. 158 <=> v2_xxreal_0(sK0) [avatar definition]
1396. v2_xxreal_0(sK0) <- (158) [avatar component clause 1394] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1398. 159 <=> v3_xxreal_0(sK0) [avatar definition]
1557. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X0) [factoring 400] {a:1,w:7,wCS:2800,sine:255,thAx:0,allAx:1,thDist:-1}
1558. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) [duplicate literal removal 1557] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
1585. 179 <=> v1_xreal_0(k7_xcmplx_0(sK0,3)) [avatar definition]
1586. v1_xreal_0(k7_xcmplx_0(sK0,3)) <- (179) [avatar component clause 1585] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1587. ~v1_xreal_0(k7_xcmplx_0(sK0,3)) <- (~179) [avatar component clause 1585] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1589. 180 <=> v3_xxreal_0(k7_xcmplx_0(sK0,3)) [avatar definition]
1591. v3_xxreal_0(k7_xcmplx_0(sK0,3)) <- (180) [avatar component clause 1589] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1597. 182 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1598. v3_xxreal_0(k6_numbers) <- (182) [avatar component clause 1597] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1602. ~v1_xreal_0(3) | ~v1_xreal_0(sK0) <- (~179) [resolution 1587,395] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1604. ~v1_xreal_0(sK0) <- (120, ~179) [subsumption resolution 1602,1096] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1605. $false <- (3, 120, ~179) [subsumption resolution 1604,449] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1606. ~3 | ~120 | 179 [avatar contradiction clause 1605]
1633. 185 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1634. ~v2_xxreal_0(k6_numbers) <- (~185) [avatar component clause 1633] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1635. v2_xxreal_0(k6_numbers) <- (185) [avatar component clause 1633] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1642. v3_xxreal_0(k6_numbers) | v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (~2) [resolution 339,444] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1660. 188 <=> v1_xboole_0(k7_xcmplx_0(sK0,3)) [avatar definition]
1661. v1_xboole_0(k7_xcmplx_0(sK0,3)) <- (188) [avatar component clause 1660] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1723. ~v1_xreal_0(3) | v3_xxreal_0(3) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (180) [resolution 381,1591] {a:3,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1732. v3_xxreal_0(3) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (120, 180) [subsumption resolution 1723,1096] {a:3,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1733. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (120, ~138, 180) [subsumption resolution 1732,1229] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1734. v3_xxreal_0(sK0) <- (3, 120, ~138, 180) [subsumption resolution 1733,449] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1742. v3_xxreal_0(k6_numbers) | v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (~2, 3) [subsumption resolution 1642,449] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1743. ~128 | 158 | 182 | 2 | ~3 [avatar split clause 1742,447,442,1597,1394,1141]
1746. 159 | ~3 | ~120 | 138 | ~180 [avatar split clause 1734,1589,1227,1094,447,1398]
1760. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (185) [resolution 1635,364] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1763. ~v1_xxreal_0(k6_numbers) <- (77, 185) [subsumption resolution 1760,817] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1764. ~126 | ~77 | ~185 [avatar split clause 1763,1633,815,1131]
1817. ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (158) [resolution 1396,363] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1818. ~v3_xxreal_0(sK0) <- (99, 158) [subsumption resolution 1817,964] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1819. ~159 | ~99 | ~158 [avatar split clause 1818,1394,962,1398]
1914. v2_xxreal_0(k6_numbers) | v1_xboole_0(k7_xcmplx_0(sK0,3)) | v3_xxreal_0(k7_xcmplx_0(sK0,3)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,3)) <- (1) [resolution 336,439] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1929. v1_xboole_0(k7_xcmplx_0(sK0,3)) | v3_xxreal_0(k7_xcmplx_0(sK0,3)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,3)) <- (1, ~185) [subsumption resolution 1914,1634] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2044. ~v1_xcmplx_0(3) | v1_xboole_0(3) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (188) [resolution 1661,389] {a:4,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2049. v1_xboole_0(3) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (123, 188) [subsumption resolution 2044,1114] {a:4,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2050. ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (~5, 123, 188) [subsumption resolution 2049,459] {a:4,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2051. v1_xboole_0(sK0) <- (~5, 97, 123, 188) [subsumption resolution 2050,945] {a:4,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2052. $false <- (~5, 97, 123, ~134, 188) [subsumption resolution 2051,1202] {a:4,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2053. 5 | ~97 | ~123 | 134 | ~188 [avatar contradiction clause 2052]
2057. v1_xboole_0(k7_xcmplx_0(sK0,3)) | v3_xxreal_0(k7_xcmplx_0(sK0,3)) | ~v1_xreal_0(k6_numbers) <- (1, 179, ~185) [subsumption resolution 1929,1586] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2058. ~128 | 180 | 188 | ~1 | ~179 | 185 [avatar split clause 2057,1633,1585,437,1660,1589,1141]
2073. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (182) [resolution 1598,361] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2075. ~v1_xxreal_0(k6_numbers) <- (77, 182) [subsumption resolution 2073,817] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2076. ~126 | ~77 | ~182 [avatar split clause 2075,1597,815,1131]
2077. 43 [avatar split clause 1558,645]
2081. $false [avatar sat refutation 440,445,450,460,470,485,490,605,620,818,824,946,965,1074,1097,1115,1134,1144,1203,1230,1362,1606,1743,1746,1764,1819,2053,2058,2076,2077] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1165,thDist:-1165}
% SZS output end Proof for xreal_1__t222_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 125
% >>> Preprocessing
% Introduced skolems: 19
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 153
% Generated clauses: 990
% Active clauses: 203
% Passive clauses: 338
% Final active clauses: 185
% Final passive clauses: 108
% Main loop iterations started: 219
% 
% >>> Simplifying Inferences
% Duplicate literals: 99
% Fw subsumption resolutions: 107
% Fw demodulations: 10
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 70
% Equational tautologies: 1
% Forward subsumptions: 159
% Backward subsumptions: 2
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Binary resolution: 241
% Unit resulting resolution: 3
% Factoring: 1
% Forward superposition: 132
% Backward superposition: 37
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 47
% Split components: 120
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 240
% SAT solver unit clauses: 78
% SAT solver binary clauses: 85
% 
% Memory used [KB]: 6268
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% congruence closure: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s
% superposition: 0.001 s
% other: 0.009 s ( own 0.004 s ) 

