i: [1,72,0,2,1,1,0,0] t28_algspec1
a: [3,115,0,2,1,72]
s: 115
i: [1,71,0,2,1,1,0,0] t28_algspec1
a: [3,120,0,2,1,71]
s: 120
i: [1,70,0,2,1,1,0,0] t28_algspec1
a: [3,125,0,2,1,70]
s: 125
i: [1,69,0,2,1,1,0,0] t28_algspec1
a: [3,130,0,2,1,69]
s: 130
i: [1,74,0,3,1,1,0,0] t28_algspec1
a: [3,105,0,3,1,74]
s: 105
i: [1,73,0,3,1,1,0,0] t28_algspec1
a: [3,110,0,3,1,73]
s: 110
i: [1,75,0,6,2,1,0,0] t28_algspec1
s: 75
i: [1,101,0,2,1,0,0,1] rc1_funct_1
a: [3,146,0,2,1,101]
s: 146
i: [1,100,0,2,1,0,0,1] rc1_funct_1
a: [3,151,0,2,1,100]
s: 151
i: [1,93,0,3,1,0,0,1] reflexivity_r1_tarski
s: 93
i: [1,92,0,4,1,0,0,255] existence_m1_subset_1
s: 92
i: [1,90,0,7,2,0,0,255] t3_subset
s: 90
i: [1,91,0,7,2,0,0,255] t3_subset
s: 91
d: [2,153,1,5,1,0,40,90,92]
s: 153
i: [1,94,0,11,5,0,0,255] reflexivity_r1_partfun1
a: [3,137,0,7,3,94]
s: 137
i: [1,76,0,8,3,0,0,255] cc2_relat_1
s: 76
d: [2,155,1,6,2,0,40,76,92]
s: 155
i: [1,77,0,8,3,0,0,255] dt_o_1_0_algspec1
s: 77
d: [2,157,1,8,3,0,40,77,75]
d: [2,158,1,6,2,1,21,157,115]
d: [2,159,1,4,1,2,21,158,120]
a: [3,163,1,4,1,159]
s: 163
i: [1,89,0,8,3,0,0,1] existence_m1_algspec1
s: 89
i: [1,78,0,9,4,0,0,1] dt_m1_algspec1
s: 78
d: [2,173,1,6,3,1,40,78,105]
d: [2,178,1,4,2,2,21,173,125]
d: [2,179,1,2,1,3,21,178,130]
a: [3,183,1,2,1,179]
s: 183
i: [1,79,0,9,4,0,0,1] dt_m1_algspec1
s: 79
d: [2,185,1,6,3,1,40,79,105]
d: [2,190,1,4,2,2,21,185,125]
d: [2,191,1,2,1,3,21,190,130]
a: [3,195,1,2,1,191]
s: 195
i: [1,95,0,13,5,0,0,255] idempotence_k1_funct_4
a: [3,141,0,9,3,95]
s: 141
i: [1,80,0,10,4,0,0,255] cc3_funct_1
s: 80
d: [2,165,1,8,3,0,40,89,75]
d: [2,166,1,6,2,1,21,165,115]
d: [2,167,1,4,1,2,21,166,120]
a: [3,171,1,4,1,167]
s: 171
i: [1,97,0,12,5,0,0,2] dt_k1_funct_4
s: 97
d: [2,156,1,7,3,0,40,76,91]
s: 156
i: [1,98,0,12,5,0,0,2] dt_k1_funct_4
s: 98
i: [1,81,0,13,5,0,0,1] t25_funct_4
s: 81
d: [2,175,1,11,5,0,40,78,89]
d: [2,176,1,7,3,0,18,175]
s: 176
i: [1,99,0,14,6,0,0,255] symmetry_r1_partfun1
s: 99
d: [2,174,1,11,5,0,40,78,77]
d: [2,177,1,7,3,0,18,174]
s: 177
i: [1,84,0,15,6,0,0,255] t52_partfun1
s: 84
i: [1,85,0,15,6,0,0,255] t52_partfun1
s: 85
d: [2,187,1,11,5,0,40,79,89]
d: [2,188,1,7,3,0,18,187]
s: 188
i: [1,83,0,16,6,0,0,255] t98_funct_4
s: 83
d: [2,186,1,11,5,0,40,79,77]
d: [2,189,1,7,3,0,18,186]
s: 189
d: [2,211,0,14,5,0,21,83,156]
s: 211
i: [1,86,0,16,6,0,0,255] t52_partfun1
s: 86
d: [2,197,1,8,3,0,40,80,92]
s: 197
i: [1,87,0,16,6,0,0,255] t52_partfun1
s: 87
d: [2,198,1,9,4,0,40,80,91]
s: 198
i: [1,96,0,17,5,0,0,255] rd3_funct_4
s: 96
i: [1,88,0,21,9,0,0,255] t52_partfun1
s: 88
d: [2,238,0,19,8,0,21,88,156]
d: [2,239,0,17,7,0,21,238,156]
d: [2,240,0,15,6,0,21,239,198]
d: [2,241,0,13,5,0,21,240,198]
s: 241
i: [1,82,0,24,8,0,0,255] t27_algspec1
s: 82
d: [2,242,1,10,4,1,40,241,110]
d: [2,248,1,8,3,2,21,242,115]
d: [2,249,1,6,2,3,21,248,120]
s: 249
d: [2,243,1,10,4,0,40,241,93]
s: 243
d: [2,261,1,26,10,1,44,82,141]
d: [2,266,1,18,6,1,18,261]
d: [2,278,1,15,5,2,21,266,137]
s: 278
d: [2,283,2,3,1,3,40,249,110]
a: [3,289,2,3,1,283]
s: 289
d: [2,244,1,24,8,0,40,241,81]
d: [2,250,1,20,7,0,21,244,98]
d: [2,251,1,16,6,0,21,250,97]
s: 251
d: [2,284,2,3,1,3,40,249,93]
a: [3,294,2,3,1,284]
s: 294
d: [2,245,1,27,9,0,40,241,86]
d: [2,252,1,23,8,0,21,245,85]
d: [2,253,1,19,7,0,21,252,84]
s: 253
d: [2,246,1,27,9,0,40,241,87]
d: [2,254,1,23,8,0,21,246,85]
d: [2,255,1,19,7,0,21,254,84]
s: 255
d: [2,301,2,7,3,1,40,243,110]
d: [2,307,2,5,2,2,21,301,115]
d: [2,308,2,3,1,3,21,307,120]
a: [3,312,2,3,1,308]
s: 312
d: [2,256,1,29,10,0,40,82,79]
d: [2,267,1,25,9,0,21,256,98]
d: [2,268,1,21,8,0,21,267,97]
s: 268
d: [2,285,2,5,1,3,40,249,153]
a: [3,299,2,5,1,285]
s: 299
d: [2,257,1,29,10,0,40,82,78]
d: [2,269,1,25,9,0,21,257,97]
d: [2,270,1,21,8,0,21,269,98]
s: 270
d: [2,262,1,31,12,0,44,82,211]
d: [2,265,1,25,9,0,18,262]
d: [2,279,1,23,8,0,21,265,156]
d: [2,280,1,21,7,0,21,279,198]
s: 280
d: [2,365,3,17,5,1,40,299,99]
d: [2,366,3,15,4,2,21,365,125]
d: [2,367,3,13,3,3,21,366,130]
a: [3,371,3,4,1,367]
s: 371
d: [2,394,4,0,0,2,52,120,371,155]
e: 394
d: [2,395,4,0,0,2,52,120,153,371,156]
e: 395
d: [2,396,4,0,0,2,52,120,92,371,76]
e: 396
d: [2,397,4,2,1,1,40,371,155]
d: [2,401,4,0,0,2,21,397,120]
e: 401
r: 371
d: [2,258,1,26,10,1,44,82,141]
d: [2,271,1,24,9,1,21,258,78]
d: [2,272,1,22,8,1,21,271,79]
s: 272
a: [3,370,3,4,1,367]
s: 370
d: [2,259,1,31,12,0,44,82,211]
d: [2,273,1,29,11,0,21,259,198]
d: [2,274,1,27,10,0,21,273,78]
d: [2,275,1,25,9,0,21,274,79]
s: 275
d: [2,263,1,40,12,0,44,82,96]
d: [2,264,1,36,10,0,18,263]
d: [2,281,1,32,9,0,21,264,98]
d: [2,282,1,28,8,0,21,281,97]
s: 282
a: [3,375,3,4,1,367]
s: 375
d: [2,430,4,0,0,3,52,120,115,375,197]
e: 430
d: [2,431,4,0,0,3,52,115,120,153,375,198]
e: 431
d: [2,432,4,0,0,3,52,120,115,92,375,80]
e: 432
d: [2,433,4,4,2,1,40,375,197]
d: [2,437,4,2,1,2,21,433,115]
d: [2,438,4,0,0,3,21,437,120]
e: 438
r: 375
d: [2,260,1,34,12,0,44,82,96]
d: [2,276,1,32,11,0,21,260,78]
d: [2,277,1,30,10,0,21,276,79]
s: 277
a: [3,374,3,4,1,367]
s: 374
d: [2,320,2,15,5,2,40,278,75]
d: [2,329,2,13,4,3,21,320,115]
d: [2,330,2,11,3,4,21,329,120]
s: 330
d: [2,384,2,21,7,0,40,280,75]
d: [2,392,2,19,6,1,21,384,115]
d: [2,393,2,17,5,2,21,392,120]
s: 393
d: [2,453,3,14,6,4,40,330,81]
d: [2,458,3,12,5,5,21,453,195]
d: [2,459,3,10,4,6,21,458,183]
a: [3,463,3,3,1,459]
s: 463
d: [2,306,2,9,3,0,40,243,153]
s: 306
a: [3,379,3,5,1,367]
s: 379
d: [2,247,2,12,4,0,40,241,153]
s: 247
d: [2,303,2,21,7,0,40,243,81]
d: [2,314,2,17,6,0,21,303,98]
d: [2,315,2,13,5,0,21,314,97]
s: 315
a: [3,466,3,7,3,459]
s: 466
d: [2,334,2,13,5,0,40,251,93]
s: 334
d: [2,500,4,10,4,1,40,466,77]
d: [2,506,4,8,3,2,21,500,115]
d: [2,507,4,6,2,3,21,506,120]
a: [3,511,4,3,1,507]
s: 511
d: [2,538,5,0,0,3,52,120,115,511,189]
e: 538
d: [2,539,5,4,2,1,40,511,189]
d: [2,541,5,2,1,2,21,539,115]
d: [2,542,5,0,0,3,21,541,120]
e: 542
r: 511
d: [2,322,2,18,7,2,40,278,78]
d: [2,327,2,14,5,2,18,322]
s: 327
d: [2,321,2,18,7,2,40,278,79]
d: [2,328,2,14,5,2,18,321]
s: 328
a: [3,510,4,3,1,507]
s: 510
d: [2,335,2,15,5,0,40,251,153]
s: 335
a: [3,515,4,3,1,507]
s: 515
d: [2,562,5,0,0,3,52,120,115,515,177]
e: 562
d: [2,563,5,4,2,1,40,515,177]
d: [2,565,5,2,1,2,21,563,115]
d: [2,566,5,0,0,3,21,565,120]
e: 566
r: 466
r: 515
d: [2,304,2,24,8,0,40,243,86]
d: [2,316,2,20,7,0,21,304,85]
d: [2,317,2,16,6,0,21,316,84]
s: 317
d: [2,305,2,24,8,0,40,243,87]
d: [2,318,2,20,7,0,21,305,85]
d: [2,319,2,16,6,0,21,318,84]
s: 319
a: [3,514,4,3,1,507]
s: 514
d: [2,351,2,16,6,0,40,253,93]
s: 351
d: [2,501,4,10,4,1,40,466,89]
d: [2,517,4,8,3,2,21,501,115]
d: [2,518,4,6,2,3,21,517,120]
a: [3,526,4,3,1,518]
s: 526
d: [2,578,5,0,0,3,52,120,115,526,176]
e: 578
d: [2,579,5,4,2,1,40,526,176]
d: [2,581,5,2,1,2,21,579,115]
d: [2,582,5,0,0,3,21,581,120]
e: 582
r: 526
d: [2,357,2,16,6,0,40,255,93]
s: 357
d: [2,352,2,18,6,0,40,253,153]
s: 352
a: [3,525,4,3,1,518]
s: 525
d: [2,358,2,18,6,0,40,255,153]
s: 358
d: [2,468,3,20,8,2,40,393,81]
d: [2,472,3,18,7,3,21,468,195]
d: [2,473,3,16,6,4,21,472,183]
a: [3,476,3,9,3,473]
s: 476
d: [2,590,4,0,0,4,52,110,105,294,476]
e: 590
d: [2,591,4,6,2,2,40,476,294]
d: [2,598,4,3,1,3,21,591,105]
d: [2,599,4,0,0,4,21,598,110]
e: 599
f: 599,566,542,507,473,284,191,179,69,70,71,72,73,74
% Refutation found. Thanks to Tanya!
% SZS status Theorem for algspec1__t28_algspec1
% SZS output start Proof for algspec1__t28_algspec1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) => ! [X2] : (m1_algspec1(X2,X0) => ? [X3] : (r1_tarski(X2,X3) & m1_algspec1(X3,X1)))))) [input t28_algspec1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) => ! [X2] : (m1_algspec1(X2,X0) => ? [X3] : (r1_tarski(X2,X3) & m1_algspec1(X3,X1)))))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
4. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_funct_1(X1))) [input cc3_funct_1]
7. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : (m1_algspec1(X1,X0) => (v1_funct_1(X1) & v1_relat_1(X1)))) [input dt_m1_algspec1]
9. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => m1_algspec1(o_1_0_algspec1(X0),X0)) [input dt_o_1_0_algspec1]
16. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
18. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => r1_tarski(X0,k1_funct_4(X1,X0)))) [input t25_funct_4]
19. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) => ! [X2] : (m1_algspec1(X2,X0) => ! [X3] : (m1_algspec1(X3,X1) => m1_algspec1(k1_funct_4(X2,X3),k1_funct_4(X0,X1))))))) [input t27_algspec1]
20. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
21. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) <=> ? [X2] : (r1_tarski(X1,X2) & r1_tarski(X0,X2) & v1_funct_1(X2) & v1_relat_1(X2))))) [input t52_partfun1]
22. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) => k1_funct_4(X1,X0) = X1))) [input t98_funct_4]
23. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 20]
24. ! [X0] : r1_tarski(X0,X0) [rectify 16]
25. ? [X0] : (? [X1] : ((? [X2] : (! [X3] : (~r1_tarski(X2,X3) | ~m1_algspec1(X3,X1)) & m1_algspec1(X2,X0)) & r1_tarski(X0,X1)) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
26. ? [X0] : (? [X1] : (? [X2] : (! [X3] : (~r1_tarski(X2,X3) | ~m1_algspec1(X3,X1)) & m1_algspec1(X2,X0)) & r1_tarski(X0,X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 25]
27. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
28. ! [X0] : (m1_algspec1(o_1_0_algspec1(X0),X0) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 9]
29. ! [X0] : (m1_algspec1(o_1_0_algspec1(X0),X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 28]
30. ! [X0] : (! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) | ~m1_algspec1(X1,X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 7]
31. ! [X0] : (! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) | ~m1_algspec1(X1,X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 30]
32. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 4]
33. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 32]
34. ! [X0] : (! [X1] : (r1_tarski(X0,k1_funct_4(X1,X0)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 18]
35. ! [X0] : (! [X1] : (r1_tarski(X0,k1_funct_4(X1,X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 34]
36. ! [X0] : (! [X1] : ((! [X2] : (! [X3] : (m1_algspec1(k1_funct_4(X2,X3),k1_funct_4(X0,X1)) | ~m1_algspec1(X3,X1)) | ~m1_algspec1(X2,X0)) | ~r1_partfun1(X0,X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 19]
37. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (m1_algspec1(k1_funct_4(X2,X3),k1_funct_4(X0,X1)) | ~m1_algspec1(X3,X1)) | ~m1_algspec1(X2,X0)) | ~r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 36]
38. ! [X0] : (! [X1] : ((k1_funct_4(X1,X0) = X1 | ~r1_tarski(X0,X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 22]
39. ! [X0] : (! [X1] : (k1_funct_4(X1,X0) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 38]
40. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> ? [X2] : (r1_tarski(X1,X2) & r1_tarski(X0,X2) & v1_funct_1(X2) & v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 21]
41. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> ? [X2] : (r1_tarski(X1,X2) & r1_tarski(X0,X2) & v1_funct_1(X2) & v1_relat_1(X2))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 40]
54. ? [X0] : (? [X1] : (? [X2] : (! [X3] : (~r1_tarski(X2,X3) | ~m1_algspec1(X3,X1)) & m1_algspec1(X2,X0)) & r1_tarski(X0,X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (! [X3] : (~r1_tarski(X2,X3) | ~m1_algspec1(X3,X1)) & m1_algspec1(X2,sK0)) & r1_tarski(sK0,X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
55. ? [X1] : (? [X2] : (! [X3] : (~r1_tarski(X2,X3) | ~m1_algspec1(X3,X1)) & m1_algspec1(X2,sK0)) & r1_tarski(sK0,X1) & v1_funct_1(X1) & v1_relat_1(X1)) => (? [X2] : (! [X3] : (~r1_tarski(X2,X3) | ~m1_algspec1(X3,sK1)) & m1_algspec1(X2,sK0)) & r1_tarski(sK0,sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
56. ? [X2] : (! [X3] : (~r1_tarski(X2,X3) | ~m1_algspec1(X3,sK1)) & m1_algspec1(X2,sK0)) => (! [X3] : (~r1_tarski(sK2,X3) | ~m1_algspec1(X3,sK1)) & m1_algspec1(sK2,sK0)) [choice axiom]
57. ((! [X3] : (~r1_tarski(sK2,X3) | ~m1_algspec1(X3,sK1)) & m1_algspec1(sK2,sK0)) & r1_tarski(sK0,sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 26,56,55,54]
58. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ! [X2] : (~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X2] : (r1_tarski(X1,X2) & r1_tarski(X0,X2) & v1_funct_1(X2) & v1_relat_1(X2)) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 41]
59. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ! [X2] : (~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X3] : (r1_tarski(X1,X3) & r1_tarski(X0,X3) & v1_funct_1(X3) & v1_relat_1(X3)) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [rectify 58]
60. ! [X1,X0] : (? [X3] : (r1_tarski(X1,X3) & r1_tarski(X0,X3) & v1_funct_1(X3) & v1_relat_1(X3)) => (r1_tarski(X1,sK3(X0,X1)) & r1_tarski(X0,sK3(X0,X1)) & v1_funct_1(sK3(X0,X1)) & v1_relat_1(sK3(X0,X1)))) [choice axiom]
61. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ! [X2] : (~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & ((r1_tarski(X1,sK3(X0,X1)) & r1_tarski(X0,sK3(X0,X1)) & v1_funct_1(sK3(X0,X1)) & v1_relat_1(sK3(X0,X1))) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [skolemisation 59,60]
64. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 23]
69. v1_relat_1(sK0) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
70. v1_funct_1(sK0) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
71. v1_relat_1(sK1) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
72. v1_funct_1(sK1) [cnf transformation 57] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
73. r1_tarski(sK0,sK1) [cnf transformation 57] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
74. m1_algspec1(sK2,sK0) [cnf transformation 57] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
75. ~m1_algspec1(X3,sK1) | ~r1_tarski(sK2,X3) [cnf transformation 57] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
76. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
77. m1_algspec1(o_1_0_algspec1(X0),X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
78. ~m1_algspec1(X1,X0) | v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. ~m1_algspec1(X1,X0) | v1_funct_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
80. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_funct_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
81. r1_tarski(X0,k1_funct_4(X1,X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
82. m1_algspec1(k1_funct_4(X2,X3),k1_funct_4(X0,X1)) | ~m1_algspec1(X3,X1) | ~m1_algspec1(X2,X0) | ~r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 37] {a:0,w:24,wCS:9600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
83. k1_funct_4(X1,X0) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 39] {a:0,w:16,wCS:6400,sine:255,thAx:0,allAx:1,thDist:-1}
88. r1_partfun1(X0,X1) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 61] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:2,thDist:-2}
91. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 64] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
93. r1_tarski(X0,X0) [cnf transformation 24] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
103. 1 <=> m1_algspec1(sK2,sK0) [avatar definition]
105. m1_algspec1(sK2,sK0) <- (1) [avatar component clause 103] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
106. 1 [avatar split clause 74,103]
108. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
110. r1_tarski(sK0,sK1) <- (2) [avatar component clause 108] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
111. 2 [avatar split clause 73,108]
113. 3 <=> v1_funct_1(sK1) [avatar definition]
115. v1_funct_1(sK1) <- (3) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
116. 3 [avatar split clause 72,113]
118. 4 <=> v1_relat_1(sK1) [avatar definition]
120. v1_relat_1(sK1) <- (4) [avatar component clause 118] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
121. 4 [avatar split clause 71,118]
123. 5 <=> v1_funct_1(sK0) [avatar definition]
125. v1_funct_1(sK0) <- (5) [avatar component clause 123] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
126. 5 [avatar split clause 70,123]
128. 6 <=> v1_relat_1(sK0) [avatar definition]
130. v1_relat_1(sK0) <- (6) [avatar component clause 128] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
131. 6 [avatar split clause 69,128]
156. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 76,91] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
173. v1_relat_1(sK2) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (1) [resolution 78,105] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
174. v1_relat_1(o_1_0_algspec1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [resolution 78,77] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
177. v1_relat_1(o_1_0_algspec1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [duplicate literal removal 174] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
178. v1_relat_1(sK2) | ~v1_relat_1(sK0) <- (1, 5) [subsumption resolution 173,125] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
179. v1_relat_1(sK2) <- (1, 5, 6) [subsumption resolution 178,130] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
181. 14 <=> v1_relat_1(sK2) [avatar definition]
183. v1_relat_1(sK2) <- (14) [avatar component clause 181] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
184. 14 | ~1 | ~5 | ~6 [avatar split clause 179,128,123,103,181]
185. v1_funct_1(sK2) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (1) [resolution 79,105] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
186. v1_funct_1(o_1_0_algspec1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [resolution 79,77] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
189. v1_funct_1(o_1_0_algspec1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [duplicate literal removal 186] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
190. v1_funct_1(sK2) | ~v1_relat_1(sK0) <- (1, 5) [subsumption resolution 185,125] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
191. v1_funct_1(sK2) <- (1, 5, 6) [subsumption resolution 190,130] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
193. 15 <=> v1_funct_1(sK2) [avatar definition]
195. v1_funct_1(sK2) <- (15) [avatar component clause 193] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
196. 15 | ~1 | ~5 | ~6 [avatar split clause 191,128,123,103,193]
198. ~r1_tarski(X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | v1_funct_1(X1) [resolution 80,91] {a:1,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
211. k1_funct_4(X1,X0) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) [subsumption resolution 83,156] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
238. r1_partfun1(X0,X1) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 88,156] {a:0,w:19,wCS:7600,sine:255,thAx:0,allAx:2,thDist:-2}
239. r1_partfun1(X0,X1) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_funct_1(X0) [subsumption resolution 238,156] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:2,thDist:-2}
240. r1_partfun1(X0,X1) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X0) [subsumption resolution 239,198] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:2,thDist:-2}
241. ~r1_tarski(X1,X2) | r1_partfun1(X0,X1) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [subsumption resolution 240,198] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
242. r1_partfun1(X0,sK0) | ~r1_tarski(X0,sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2) [resolution 241,110] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
248. r1_partfun1(X0,sK0) | ~r1_tarski(X0,sK1) | ~v1_relat_1(sK1) <- (2, 3) [subsumption resolution 242,115] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
249. ~r1_tarski(X0,sK1) | r1_partfun1(X0,sK0) <- (2, 3, 4) [subsumption resolution 248,120] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
262. m1_algspec1(k1_funct_4(X5,X6),X3) | ~m1_algspec1(X6,X4) | ~m1_algspec1(X5,X3) | ~r1_partfun1(X3,X4) | ~v1_funct_1(X4) | ~v1_relat_1(X4) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~r1_tarski(X4,X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_funct_1(X4) [superposition 82,211] {a:1,w:31,wCS:12400,sine:255,thAx:0,allAx:2,thDist:-2}
265. m1_algspec1(k1_funct_4(X5,X6),X3) | ~m1_algspec1(X6,X4) | ~m1_algspec1(X5,X3) | ~r1_partfun1(X3,X4) | ~v1_funct_1(X4) | ~v1_relat_1(X4) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~r1_tarski(X4,X3) [duplicate literal removal 262] {a:1,w:25,wCS:10000,sine:255,thAx:0,allAx:2,thDist:-2}
279. m1_algspec1(k1_funct_4(X5,X6),X3) | ~m1_algspec1(X6,X4) | ~m1_algspec1(X5,X3) | ~r1_partfun1(X3,X4) | ~v1_funct_1(X4) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~r1_tarski(X4,X3) [subsumption resolution 265,156] {a:1,w:23,wCS:9200,sine:255,thAx:0,allAx:2,thDist:-2}
280. m1_algspec1(k1_funct_4(X5,X6),X3) | ~m1_algspec1(X6,X4) | ~m1_algspec1(X5,X3) | ~r1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~r1_tarski(X4,X3) [subsumption resolution 279,198] {a:1,w:21,wCS:8400,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
284. r1_partfun1(sK1,sK0) <- (2, 3, 4) [resolution 249,93] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
292. 17 <=> r1_partfun1(sK1,sK0) [avatar definition]
294. r1_partfun1(sK1,sK0) <- (17) [avatar component clause 292] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
295. 17 | ~2 | ~3 | ~4 [avatar split clause 284,118,113,108,292]
384. ~m1_algspec1(X0,X1) | ~m1_algspec1(X2,sK1) | ~r1_partfun1(sK1,X1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~r1_tarski(X1,sK1) | ~r1_tarski(sK2,k1_funct_4(X2,X0)) [resolution 280,75] {a:2,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
392. ~m1_algspec1(X0,X1) | ~m1_algspec1(X2,sK1) | ~r1_partfun1(sK1,X1) | ~v1_relat_1(sK1) | ~r1_tarski(X1,sK1) | ~r1_tarski(sK2,k1_funct_4(X2,X0)) <- (3) [subsumption resolution 384,115] {a:2,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
393. ~r1_tarski(sK2,k1_funct_4(X2,X0)) | ~m1_algspec1(X2,sK1) | ~r1_partfun1(sK1,X1) | ~r1_tarski(X1,sK1) | ~m1_algspec1(X0,X1) <- (3, 4) [subsumption resolution 392,120] {a:2,w:17,wCS:1700,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
465. 24 <=> ! [X0] : (~m1_algspec1(X0,sK1) | ~v1_relat_1(X0) | ~v1_funct_1(X0)) [avatar definition]
466. ~m1_algspec1(X0,sK1) | ~v1_relat_1(X0) | ~v1_funct_1(X0) <- (24) [avatar component clause 465] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
468. ~m1_algspec1(X0,sK1) | ~r1_partfun1(sK1,X1) | ~r1_tarski(X1,sK1) | ~m1_algspec1(sK2,X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (3, 4) [resolution 393,81] {a:3,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
472. ~m1_algspec1(X0,sK1) | ~r1_partfun1(sK1,X1) | ~r1_tarski(X1,sK1) | ~m1_algspec1(sK2,X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_relat_1(sK2) <- (3, 4, 15) [subsumption resolution 468,195] {a:3,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
473. ~m1_algspec1(X0,sK1) | ~r1_partfun1(sK1,X1) | ~r1_tarski(X1,sK1) | ~m1_algspec1(sK2,X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (3, 4, 14, 15) [subsumption resolution 472,183] {a:3,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
475. 25 <=> ! [X1] : (~r1_partfun1(sK1,X1) | ~m1_algspec1(sK2,X1) | ~r1_tarski(X1,sK1)) [avatar definition]
476. ~r1_partfun1(sK1,X1) | ~m1_algspec1(sK2,X1) | ~r1_tarski(X1,sK1) <- (25) [avatar component clause 475] {a:3,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
477. 25 | 24 | ~3 | ~4 | ~14 | ~15 [avatar split clause 473,193,181,118,113,465,475]
500. ~v1_relat_1(o_1_0_algspec1(sK1)) | ~v1_funct_1(o_1_0_algspec1(sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (24) [resolution 466,77] {a:4,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
506. ~v1_relat_1(o_1_0_algspec1(sK1)) | ~v1_funct_1(o_1_0_algspec1(sK1)) | ~v1_relat_1(sK1) <- (3, 24) [subsumption resolution 500,115] {a:4,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
507. ~v1_relat_1(o_1_0_algspec1(sK1)) | ~v1_funct_1(o_1_0_algspec1(sK1)) <- (3, 4, 24) [subsumption resolution 506,120] {a:4,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
509. 26 <=> v1_funct_1(o_1_0_algspec1(sK1)) [avatar definition]
511. ~v1_funct_1(o_1_0_algspec1(sK1)) <- (~26) [avatar component clause 509] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
513. 27 <=> v1_relat_1(o_1_0_algspec1(sK1)) [avatar definition]
515. ~v1_relat_1(o_1_0_algspec1(sK1)) <- (~27) [avatar component clause 513] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
516. ~26 | ~27 | ~3 | ~4 | ~24 [avatar split clause 507,465,118,113,513,509]
539. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~26) [resolution 511,189] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
541. ~v1_relat_1(sK1) <- (3, ~26) [subsumption resolution 539,115] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
542. $false <- (3, 4, ~26) [subsumption resolution 541,120] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
543. ~3 | ~4 | 26 [avatar contradiction clause 542]
563. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~27) [resolution 515,177] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
565. ~v1_relat_1(sK1) <- (3, ~27) [subsumption resolution 563,115] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
566. $false <- (3, 4, ~27) [subsumption resolution 565,120] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
567. ~3 | ~4 | 27 [avatar contradiction clause 566]
591. ~m1_algspec1(sK2,sK0) | ~r1_tarski(sK0,sK1) <- (17, 25) [resolution 476,294] {a:4,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
598. ~r1_tarski(sK0,sK1) <- (1, 17, 25) [subsumption resolution 591,105] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
599. $false <- (1, 2, 17, 25) [subsumption resolution 598,110] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
600. ~1 | ~2 | ~17 | ~25 [avatar contradiction clause 599]
607. $false [avatar sat refutation 106,111,116,121,126,131,184,196,295,477,516,543,567,600] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:479,thDist:-479}
% SZS output end Proof for algspec1__t28_algspec1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 33
% Generated clauses: 419
% Active clauses: 92
% Passive clauses: 108
% Final active clauses: 86
% Final passive clauses: 12
% Main loop iterations started: 94
% 
% >>> Simplifying Inferences
% Duplicate literals: 192
% Fw subsumption resolutions: 100
% 
% >>> Deletion Inferences
% Simple tautologies: 43
% Forward subsumptions: 56
% 
% >>> Generating Inferences
% Binary resolution: 102
% Unit resulting resolution: 10
% Forward superposition: 65
% Backward superposition: 6
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 15
% Sat splitting refutations: 16
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 42
% SAT solver unit clauses: 8
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 5373
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% backward subsumption index maintenance: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.004 s ( own 0.001 s ) 

