i: [1,37,0,2,1,1,0,0] t3_isomichi
a: [3,66,0,2,1,37]
s: 66
i: [1,36,0,2,1,1,0,0] t3_isomichi
a: [3,71,0,2,1,36]
s: 71
i: [1,40,0,3,1,1,0,0] t3_isomichi
a: [3,51,0,3,1,40]
s: 51
i: [1,39,0,3,1,1,0,0] t3_isomichi
a: [3,56,0,3,1,39]
s: 56
i: [1,38,0,5,1,1,0,0] t3_isomichi
a: [3,61,0,5,1,38]
s: 61
i: [1,47,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,76,0,2,1,47]
s: 76
i: [1,46,0,3,1,0,0,255] reflexivity_r1_tarski
s: 46
i: [1,45,0,4,1,0,0,1] existence_m1_subset_1
s: 45
i: [1,44,0,7,2,0,0,255] t3_subset
s: 44
i: [1,43,0,8,3,0,0,1] rc6_pre_topc
s: 43
i: [1,42,0,10,3,0,0,1] rc6_pre_topc
s: 42
i: [1,41,0,15,5,0,0,1] cc1_isomichi
s: 41
d: [2,78,1,0,0,5,52,71,66,51,56,61,41]
e: 78
d: [2,80,1,10,4,1,40,41,61]
d: [2,85,1,7,3,2,21,80,56]
d: [2,86,1,4,2,3,21,85,51]
d: [2,87,1,2,1,4,21,86,66]
d: [2,88,1,0,0,5,21,87,71]
e: 88
f: 88,36,37,38,39,40
% Refutation found. Thanks to Tanya!
% SZS status Theorem for isomichi__t3_isomichi
% SZS output start Proof for isomichi__t3_isomichi
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v4_pre_topc(X1,X0) => v1_isomichi(X1,X0)))) [input t3_isomichi]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v4_pre_topc(X1,X0) => v1_isomichi(X1,X0)))) [negated conjecture 1]
12. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v4_pre_topc(X1,X0) => v1_isomichi(X1,X0)))) [input cc1_isomichi]
20. ? [X0] : (? [X1] : ((~v1_isomichi(X1,X0) & v4_pre_topc(X1,X0)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & (l1_pre_topc(X0) & v2_pre_topc(X0))) [ennf transformation 2]
21. ? [X0] : (? [X1] : (~v1_isomichi(X1,X0) & v4_pre_topc(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 20]
22. ! [X0] : (! [X1] : ((v1_isomichi(X1,X0) | ~v4_pre_topc(X1,X0)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 12]
23. ! [X0] : (! [X1] : (v1_isomichi(X1,X0) | ~v4_pre_topc(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 22]
27. ? [X0] : (? [X1] : (~v1_isomichi(X1,X0) & v4_pre_topc(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) => (? [X1] : (~v1_isomichi(X1,sK0) & v4_pre_topc(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0)) [choice axiom]
28. ? [X1] : (~v1_isomichi(X1,sK0) & v4_pre_topc(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => (~v1_isomichi(sK1,sK0) & v4_pre_topc(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
29. (~v1_isomichi(sK1,sK0) & v4_pre_topc(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) [skolemisation 21,28,27]
36. v2_pre_topc(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
37. l1_pre_topc(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 29] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. v4_pre_topc(sK1,sK0) [cnf transformation 29] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. ~v1_isomichi(sK1,sK0) [cnf transformation 29] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v4_pre_topc(X1,X0) | v1_isomichi(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 23] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. 1 <=> v1_isomichi(sK1,sK0) [avatar definition]
51. ~v1_isomichi(sK1,sK0) <- (~1) [avatar component clause 49] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
52. ~1 [avatar split clause 40,49]
54. 2 <=> v4_pre_topc(sK1,sK0) [avatar definition]
56. v4_pre_topc(sK1,sK0) <- (2) [avatar component clause 54] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
57. 2 [avatar split clause 39,54]
59. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
61. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (3) [avatar component clause 59] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
62. 3 [avatar split clause 38,59]
64. 4 <=> l1_pre_topc(sK0) [avatar definition]
66. l1_pre_topc(sK0) <- (4) [avatar component clause 64] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
67. 4 [avatar split clause 37,64]
69. 5 <=> v2_pre_topc(sK0) [avatar definition]
71. v2_pre_topc(sK0) <- (5) [avatar component clause 69] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. 5 [avatar split clause 36,69]
80. ~v4_pre_topc(sK1,sK0) | v1_isomichi(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (3) [resolution 41,61] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
85. v1_isomichi(sK1,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (2, 3) [subsumption resolution 80,56] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
86. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~1, 2, 3) [subsumption resolution 85,51] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
87. ~v2_pre_topc(sK0) <- (~1, 2, 3, 4) [subsumption resolution 86,66] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
88. $false <- (~1, 2, 3, 4, 5) [subsumption resolution 87,71] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
89. 1 | ~2 | ~3 | ~4 | ~5 [avatar contradiction clause 88]
91. $false [avatar sat refutation 52,57,62,67,72,89] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:43,thDist:-43}
% SZS output end Proof for isomichi__t3_isomichi
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 29
% Active clauses: 12
% Passive clauses: 12
% Final active clauses: 12
% Main loop iterations started: 12
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 5
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 6
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% sat proof minimization: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

