i: [1,197,0,2,1,1,0,0] t22_tops_3
a: [3,284,0,2,1,197]
s: 284
i: [1,196,0,2,1,1,0,0] t22_tops_3
a: [3,289,0,2,1,196]
s: 289
i: [1,195,0,2,1,1,0,0] t22_tops_3
a: [3,294,0,2,1,195]
s: 294
i: [1,198,0,4,1,1,0,0] t22_tops_3
a: [3,279,0,4,1,198]
s: 279
i: [1,199,0,2,1,0,0,255] spc1_boole
a: [3,299,0,2,1,199]
s: 299
i: [1,200,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,304,0,2,1,200]
s: 304
i: [1,274,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,314,0,2,1,274]
s: 314
i: [1,275,0,2,1,0,0,2] existence_l1_struct_0
a: [3,319,0,2,1,275]
s: 319
i: [1,272,0,3,1,0,0,255] reflexivity_r1_tarski
s: 272
i: [1,201,0,3,1,0,0,255] d2_xboole_0
a: [3,309,0,3,1,201]
s: 309
i: [1,202,0,4,2,0,0,1] dt_l1_pre_topc
s: 202
d: [2,321,1,2,1,1,40,202,284]
a: [3,326,1,2,1,321]
s: 326
i: [1,271,0,4,1,0,0,2] existence_m1_subset_1
s: 271
d: [2,322,1,2,1,1,40,202,314]
a: [3,331,1,2,1,322]
s: 331
i: [1,206,0,5,2,0,0,1] fc3_struct_0
s: 206
i: [1,223,0,5,2,0,0,2] t6_boole
s: 223
i: [1,268,0,5,2,0,0,255] t7_boole
s: 268
d: [2,333,0,5,2,1,22,223,309]
s: 333
i: [1,205,0,6,2,0,0,255] rc5_tops_1
s: 205
i: [1,207,0,6,2,0,0,1] d2_struct_0
s: 207
i: [1,209,0,6,3,0,0,255] cc7_struct_0
s: 209
i: [1,210,0,6,3,0,0,255] cc5_struct_0
s: 210
i: [1,212,0,6,3,0,0,255] cc2_struct_0
s: 212
d: [2,337,1,4,2,1,40,210,319]
a: [3,343,1,2,1,337]
s: 343
i: [1,213,0,6,3,0,0,255] cc17_struct_0
s: 213
i: [1,214,0,6,3,0,0,255] cc6_struct_0
s: 214
d: [2,339,2,4,2,1,40,210,331]
a: [3,352,2,2,1,339]
s: 352
i: [1,215,0,6,3,0,0,255] cc1_struct_0
s: 215
i: [1,218,0,6,3,0,0,255] cc4_struct_0
s: 218
i: [1,264,0,6,2,0,0,255] t1_subset
s: 264
i: [1,273,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 273
d: [2,336,0,6,2,1,22,207,309]
s: 336
i: [1,216,0,7,3,0,0,255] cc8_struct_0
s: 216
d: [2,371,1,5,2,2,44,279,336]
d: [2,374,1,3,1,3,21,371,326]
a: [3,378,1,3,1,374]
s: 378
i: [1,219,0,7,3,0,0,255] cc11_struct_0
s: 219
i: [1,220,0,7,3,0,0,255] cc11_struct_0
s: 220
d: [2,381,2,5,2,2,44,378,333]
s: 381
i: [1,221,0,7,3,0,0,255] cc9_struct_0
s: 221
i: [1,224,0,7,3,0,0,255] fc9_struct_0
s: 224
i: [1,225,0,7,3,0,0,255] fc6_struct_0
s: 225
i: [1,227,0,7,3,0,0,255] rc22_struct_0
s: 227
i: [1,238,0,7,3,0,0,1] fc2_struct_0
s: 238
i: [1,240,0,7,3,0,0,1] rc4_struct_0
s: 240
i: [1,242,0,7,3,0,0,255] rc21_struct_0
s: 242
i: [1,243,0,7,3,0,0,255] rc21_struct_0
s: 243
i: [1,261,0,7,3,0,0,255] fc8_struct_0
s: 261
i: [1,262,0,7,3,0,0,255] fc7_struct_0
s: 262
i: [1,263,0,7,3,0,0,1] fc1_struct_0
s: 263
i: [1,266,0,7,2,0,0,255] t3_subset
s: 266
i: [1,267,0,7,3,0,0,2] t8_boole
s: 267
d: [2,380,0,7,3,1,22,216,309]
s: 380
d: [2,382,0,7,3,1,22,221,309]
s: 382
i: [1,204,0,8,2,0,0,255] rc5_tops_1
s: 204
d: [2,334,1,7,3,1,44,333,333]
s: 334
i: [1,208,0,8,2,0,0,1] dt_k1_struct_0
s: 208
i: [1,251,0,8,3,0,0,1] rc7_tops_1
s: 251
d: [2,395,3,7,3,2,40,251,381]
d: [2,396,3,5,2,3,21,395,284]
d: [2,397,3,3,1,4,21,396,289]
a: [3,401,3,3,1,397]
s: 401
i: [1,253,0,8,3,0,0,3] rc6_pre_topc
s: 253
d: [2,393,1,9,3,1,44,208,336]
d: [2,394,1,7,2,1,18,393]
s: 394
i: [1,255,0,8,3,0,0,2] rc1_tops_1
s: 255
i: [1,257,0,8,3,0,0,2] rc10_pre_topc
s: 257
i: [1,259,0,8,3,0,0,2] rc2_tops_1
s: 259
i: [1,260,0,8,3,0,0,2] rc2_tops_1
s: 260
i: [1,265,0,8,3,0,0,255] t2_subset
s: 265
i: [1,222,0,9,4,0,0,255] cc10_struct_0
s: 222
i: [1,229,0,9,4,0,0,255] rc6_tops_1
s: 229
d: [2,408,1,6,2,0,40,265,271]
s: 408
i: [1,232,0,9,4,0,0,3] rc7_pre_topc
s: 232
d: [2,413,2,6,2,0,40,408,273]
s: 413
i: [1,235,0,9,4,0,0,2] rc3_tops_1
s: 235
i: [1,270,0,9,3,0,0,255] t5_subset
s: 270
d: [2,418,2,8,3,1,40,270,394]
a: [3,423,2,3,1,418]
s: 423
i: [1,226,0,10,3,0,0,255] rc22_struct_0
s: 226
d: [2,420,1,7,2,0,40,270,271]
s: 420
i: [1,230,0,10,4,0,0,255] rc6_tops_1
s: 230
i: [1,233,0,10,4,0,0,3] rc7_pre_topc
s: 233
d: [2,432,2,6,2,0,40,420,408]
s: 432
i: [1,236,0,10,4,0,0,2] rc3_tops_1
s: 236
d: [2,416,1,8,3,0,40,270,266]
s: 416
i: [1,237,0,10,4,0,0,2] rc3_tops_1
s: 237
i: [1,239,0,10,3,0,0,1] rc4_struct_0
s: 239
d: [2,389,1,9,4,1,44,380,333]
s: 389
i: [1,241,0,10,3,0,0,255] rc21_struct_0
s: 241
d: [2,391,1,9,4,1,44,382,333]
s: 391
i: [1,250,0,10,3,0,0,1] rc7_tops_1
s: 250
i: [1,252,0,10,3,0,0,3] rc6_pre_topc
s: 252
d: [2,417,1,9,3,0,40,270,208]
s: 417
i: [1,254,0,10,3,0,0,2] rc1_tops_1
s: 254
d: [2,419,1,9,3,0,40,270,204]
s: 419
i: [1,256,0,10,3,0,0,2] rc10_pre_topc
s: 256
i: [1,258,0,10,3,0,0,2] rc2_tops_1
s: 258
d: [2,459,2,8,3,0,40,419,408]
s: 459
i: [1,269,0,10,3,0,0,255] t4_subset
s: 269
d: [2,464,3,9,4,0,40,459,263]
d: [2,465,3,7,3,0,21,464,202]
s: 465
i: [1,203,0,12,4,0,0,255] cc2_tops_1
s: 203
i: [1,228,0,12,4,0,0,255] rc6_tops_1
s: 228
d: [2,481,2,9,4,1,40,203,394]
d: [2,499,2,7,3,2,21,481,304]
d: [2,500,2,5,2,2,21,499,202]
s: 500
i: [1,231,0,12,4,0,0,3] rc7_pre_topc
s: 231
d: [2,507,3,7,3,2,44,500,333]
s: 507
r: 203
i: [1,234,0,12,4,0,0,2] rc3_tops_1
s: 234
i: [1,244,0,14,5,0,0,1] cc3_tops_1
s: 244
d: [2,520,2,11,5,1,40,244,394]
d: [2,544,2,9,4,2,21,520,304]
d: [2,545,2,7,3,2,21,544,202]
s: 545
d: [2,547,3,0,0,5,52,289,284,378,545]
e: 547
d: [2,548,3,4,2,3,40,545,378]
d: [2,552,3,2,1,4,21,548,284]
d: [2,553,3,0,0,5,21,552,289]
e: 553
d: [2,549,3,6,3,3,40,545,381]
d: [2,555,3,4,2,4,21,549,284]
d: [2,556,3,2,1,5,21,555,289]
d: [2,557,3,0,0,5,21,556,304]
e: 557
f: 557,374,321,201,200,196,197,198
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tops_3__t22_tops_3
% SZS output start Proof for tops_3__t22_tops_3
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => v3_tops_1(k1_struct_0(X0),X0)) [input t22_tops_3]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => v3_tops_1(k1_struct_0(X0),X0)) [negated conjecture 1]
20. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
27. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v1_xboole_0(X1) => v3_tops_1(X1,X0)))) [input cc3_tops_1]
58. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
59. o_0_0_xboole_0 = k1_xboole_0 [input d2_xboole_0]
61. ! [X0] : (l1_struct_0(X0) => m1_subset_1(k1_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0)))) [input dt_k1_struct_0]
62. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
64. ! [X0] : (l1_struct_0(X0) => k1_struct_0(X0) = k1_xboole_0) [input d2_struct_0]
75. ? [X0] : (~v3_tops_1(k1_struct_0(X0),X0) & (l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
76. ? [X0] : (~v3_tops_1(k1_struct_0(X0),X0) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [flattening 75]
77. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 62]
82. ! [X0] : (k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0)) [ennf transformation 64]
83. ! [X0] : (m1_subset_1(k1_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0)) [ennf transformation 61]
106. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 58]
125. ! [X0] : (! [X1] : ((v3_tops_1(X1,X0) | ~v1_xboole_0(X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 27]
126. ! [X0] : (! [X1] : (v3_tops_1(X1,X0) | ~v1_xboole_0(X1) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 125]
163. ? [X0] : (~v3_tops_1(k1_struct_0(X0),X0) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (~v3_tops_1(k1_struct_0(sK0),sK0) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0)) [choice axiom]
164. ~v3_tops_1(k1_struct_0(sK0),sK0) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0) [skolemisation 76,163]
196. v2_pre_topc(sK0) [cnf transformation 164] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
197. l1_pre_topc(sK0) [cnf transformation 164] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
198. ~v3_tops_1(k1_struct_0(sK0),sK0) [cnf transformation 164] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
200. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 20] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
201. o_0_0_xboole_0 = k1_xboole_0 [cnf transformation 59] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
202. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 77] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
207. k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0) [cnf transformation 82] {a:0,w:6,wCS:2400,sine:1,thAx:0,allAx:1,thDist:-1}
208. m1_subset_1(k1_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) [cnf transformation 83] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
223. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 106] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
244. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v1_xboole_0(X1) | v3_tops_1(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 126] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
277. 1 <=> v3_tops_1(k1_struct_0(sK0),sK0) [avatar definition]
279. ~v3_tops_1(k1_struct_0(sK0),sK0) <- (~1) [avatar component clause 277] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
280. ~1 [avatar split clause 198,277]
282. 2 <=> l1_pre_topc(sK0) [avatar definition]
284. l1_pre_topc(sK0) <- (2) [avatar component clause 282] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
285. 2 [avatar split clause 197,282]
287. 3 <=> v2_pre_topc(sK0) [avatar definition]
289. v2_pre_topc(sK0) <- (3) [avatar component clause 287] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
290. 3 [avatar split clause 196,287]
302. 6 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
304. v1_xboole_0(o_0_0_xboole_0) <- (6) [avatar component clause 302] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
305. 6 [avatar split clause 200,302]
307. 7 <=> o_0_0_xboole_0 = k1_xboole_0 [avatar definition]
309. o_0_0_xboole_0 = k1_xboole_0 <- (7) [avatar component clause 307] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
310. 7 [avatar split clause 201,307]
321. l1_struct_0(sK0) <- (2) [resolution 202,284] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
324. 10 <=> l1_struct_0(sK0) [avatar definition]
326. l1_struct_0(sK0) <- (10) [avatar component clause 324] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
327. 10 | ~2 [avatar split clause 321,282,324]
333. o_0_0_xboole_0 = X0 | ~v1_xboole_0(X0) <- (7) [forward demodulation 223,309] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
336. k1_struct_0(X0) = o_0_0_xboole_0 | ~l1_struct_0(X0) <- (7) [forward demodulation 207,309] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
371. ~v3_tops_1(o_0_0_xboole_0,sK0) | ~l1_struct_0(sK0) <- (~1, 7) [superposition 279,336] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
374. ~v3_tops_1(o_0_0_xboole_0,sK0) <- (~1, 7, 10) [subsumption resolution 371,326] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
376. 18 <=> v3_tops_1(o_0_0_xboole_0,sK0) [avatar definition]
378. ~v3_tops_1(o_0_0_xboole_0,sK0) <- (~18) [avatar component clause 376] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
379. ~18 | 1 | ~7 | ~10 [avatar split clause 374,324,307,277,376]
381. ~v3_tops_1(X0,sK0) | ~v1_xboole_0(X0) <- (7, ~18) [superposition 378,333] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
393. m1_subset_1(o_0_0_xboole_0,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) | ~l1_struct_0(X0) <- (7) [superposition 208,336] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
394. m1_subset_1(o_0_0_xboole_0,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) <- (7) [duplicate literal removal 393] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
520. ~v1_xboole_0(o_0_0_xboole_0) | v3_tops_1(o_0_0_xboole_0,X3) | ~l1_pre_topc(X3) | ~v2_pre_topc(X3) | ~l1_struct_0(X3) <- (7) [resolution 244,394] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:3,thDist:-3}
544. v3_tops_1(o_0_0_xboole_0,X3) | ~l1_pre_topc(X3) | ~v2_pre_topc(X3) | ~l1_struct_0(X3) <- (6, 7) [subsumption resolution 520,304] {a:2,w:9,wCS:3600,sine:1,thAx:0,allAx:3,thDist:-3}
545. v3_tops_1(o_0_0_xboole_0,X3) | ~l1_pre_topc(X3) | ~v2_pre_topc(X3) <- (6, 7) [subsumption resolution 544,202] {a:2,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
549. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | ~v1_xboole_0(o_0_0_xboole_0) <- (6, 7, ~18) [resolution 545,381] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
555. ~v2_pre_topc(sK0) | ~v1_xboole_0(o_0_0_xboole_0) <- (2, 6, 7, ~18) [subsumption resolution 549,284] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
556. ~v1_xboole_0(o_0_0_xboole_0) <- (2, 3, 6, 7, ~18) [subsumption resolution 555,289] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
557. $false <- (2, 3, 6, 7, ~18) [subsumption resolution 556,304] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
558. ~2 | ~3 | ~6 | ~7 | 18 [avatar contradiction clause 557]
559. $false [avatar sat refutation 280,285,290,305,310,327,379,558] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:107,thDist:-107}
% SZS output end Proof for tops_3__t22_tops_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 63
% >>> Preprocessing
% Introduced skolems: 16
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 81
% Generated clauses: 277
% Active clauses: 93
% Passive clauses: 165
% Final active clauses: 92
% Final passive clauses: 63
% Main loop iterations started: 102
% 
% >>> Simplifying Inferences
% Duplicate literals: 44
% Fw subsumption resolutions: 17
% Fw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 32
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 112
% Unit resulting resolution: 1
% Forward superposition: 9
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 9
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% term sharing: 0.001 s
% resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.003 s ( own 0.001 s ) 

