i: [1,106,0,2,1,1,0,0] t4_bvfunc11
a: [3,169,0,2,1,106]
s: 169
i: [1,108,0,3,1,1,0,0] t4_bvfunc11
a: [3,159,0,3,1,108]
s: 159
i: [1,107,0,3,1,1,0,0] t4_bvfunc11
a: [3,164,0,3,1,107]
s: 164
i: [1,110,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,174,0,2,1,110]
s: 174
i: [1,111,0,2,1,0,0,255] fc1_xboole_0
a: [3,179,0,2,1,111]
s: 179
i: [1,145,0,2,1,0,0,1] rc2_xboole_0
a: [3,189,0,2,1,145]
s: 189
i: [1,146,0,2,1,0,0,1] rc1_xboole_0
a: [3,194,0,2,1,146]
s: 194
i: [1,109,0,20,2,1,0,0] t4_bvfunc11
a: [3,150,0,10,1,109]
s: 150
i: [1,135,0,3,1,0,0,1] reflexivity_r1_tarski
s: 135
i: [1,136,0,3,1,0,0,2] reflexivity_r1_setfam_1
s: 136
i: [1,112,0,3,1,0,0,255] d2_xboole_0
a: [3,184,0,3,1,112]
s: 184
r: 179
i: [1,113,0,4,1,0,0,1] dt_k10_eqrel_1
s: 113
i: [1,133,0,4,1,0,0,1] existence_m1_eqrel_1
s: 133
i: [1,134,0,4,1,0,0,1] existence_m1_subset_1
s: 134
i: [1,120,0,5,2,0,0,255] rc2_eqrel_1
s: 120
i: [1,122,0,5,2,0,0,255] rc4_eqrel_1
s: 122
i: [1,123,0,5,2,0,0,255] t6_boole
s: 123
i: [1,130,0,5,2,0,0,3] t7_boole
s: 130
d: [2,197,0,5,2,1,22,123,184]
s: 197
i: [1,114,0,6,2,0,0,1] dt_k6_partit1
s: 114
i: [1,119,0,6,2,0,0,255] rc2_eqrel_1
s: 119
i: [1,121,0,6,2,0,0,255] rc4_eqrel_1
s: 121
i: [1,124,0,6,2,0,0,3] t1_subset
s: 124
i: [1,137,0,6,2,0,0,3] antisymmetry_r2_hidden
s: 137
i: [1,115,0,7,3,0,0,1] cc1_eqrel_1
s: 115
i: [1,127,0,7,2,0,0,1] t3_subset
s: 127
d: [2,202,1,4,2,1,40,115,159]
d: [2,209,1,2,1,2,21,202,169]
a: [3,213,1,2,1,209]
s: 213
i: [1,128,0,7,2,0,0,1] t3_subset
s: 128
d: [2,201,1,5,2,0,40,115,113]
s: 201
i: [1,129,0,7,3,0,0,255] t8_boole
s: 129
i: [1,125,0,8,2,0,0,2] dt_m1_eqrel_1
s: 125
d: [2,205,1,5,2,0,40,115,133]
s: 205
i: [1,126,0,8,3,0,0,3] t2_subset
s: 126
d: [2,218,1,5,2,1,40,126,164]
d: [2,222,1,3,1,2,21,218,169]
a: [3,226,1,3,1,222]
s: 226
i: [1,116,0,9,3,0,0,1] t32_partit1
s: 116
i: [1,117,0,9,3,0,0,1] t32_partit1
s: 117
d: [2,228,2,3,1,1,40,226,137]
a: [3,234,2,3,1,228]
s: 234
i: [1,132,0,9,3,0,0,2] t5_subset
s: 132
d: [2,200,1,7,3,0,40,115,114]
d: [2,208,1,5,2,0,18,200]
s: 208
i: [1,131,0,10,3,0,0,2] t4_subset
s: 131
i: [1,141,0,11,3,0,0,2] existence_m1_bvfunc_1
s: 141
d: [2,215,1,5,1,0,40,127,134]
s: 215
i: [1,139,0,12,4,0,0,2] redefinition_m1_bvfunc_1
s: 139
d: [2,219,1,6,2,0,40,126,134]
s: 219
i: [1,140,0,12,4,0,0,2] redefinition_m1_bvfunc_1
s: 140
i: [1,138,0,13,4,0,0,2] dt_m1_bvfunc_1
s: 138
d: [2,244,2,6,2,0,40,219,137]
s: 244
i: [1,142,0,15,4,0,0,255] dt_k11_eqrel_1
s: 142
d: [2,198,1,7,3,1,44,197,197]
s: 198
i: [1,143,0,15,4,0,0,1] dt_k15_bvfunc_1
s: 143
i: [1,144,0,17,4,0,0,255] redefinition_k15_bvfunc_1
s: 144
d: [2,217,1,7,2,0,40,125,127]
s: 217
i: [1,118,0,23,6,0,0,2] t1_bvfunc11
s: 118
d: [2,263,1,16,5,1,40,118,150]
d: [2,264,1,13,4,2,21,263,159]
d: [2,265,1,10,3,3,21,264,164]
d: [2,266,1,8,2,4,21,265,169]
a: [3,270,1,4,1,266]
s: 270
d: [2,276,2,0,0,2,52,169,270,114]
e: 276
d: [2,277,2,2,1,1,40,270,114]
d: [2,279,2,0,0,2,21,277,169]
e: 279
r: 270
a: [3,269,1,4,1,266]
s: 269
a: [3,274,1,4,1,266]
s: 274
d: [2,282,2,0,0,3,52,169,159,274,116]
e: 282
d: [2,283,2,5,2,1,40,274,116]
d: [2,285,2,2,1,2,21,283,159]
d: [2,286,2,0,0,3,21,285,169]
e: 286
r: 150
r: 274
a: [3,273,1,4,1,266]
s: 273
a: [3,149,0,10,1,109]
s: 149
a: [3,154,0,10,1,109]
s: 154
d: [2,288,1,16,5,1,40,154,118]
d: [2,289,1,13,4,2,21,288,159]
d: [2,290,1,9,3,2,21,289,113]
d: [2,291,1,6,2,3,21,290,164]
d: [2,292,1,4,1,4,21,291,169]
a: [3,296,1,4,1,292]
s: 296
d: [2,298,2,0,0,3,52,169,159,296,117]
e: 298
d: [2,299,2,5,2,1,40,296,117]
d: [2,301,2,2,1,2,21,299,159]
d: [2,302,2,0,0,3,21,301,169]
e: 302
f: 302,292,286,279,266,106,107,108,109
% Refutation found. Thanks to Tanya!
% SZS status Theorem for bvfunc11__t4_bvfunc11
% SZS output start Proof for bvfunc11__t4_bvfunc11
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_eqrel_1(X2,X0) => (r1_tarski(k15_bvfunc_1(X0,X1,k10_eqrel_1(X0)),k15_bvfunc_1(X0,X1,X2)) & r1_tarski(k15_bvfunc_1(X0,X1,X2),k15_bvfunc_1(X0,X1,k6_partit1(X0))))))) [input t4_bvfunc11]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_eqrel_1(X2,X0) => (r1_tarski(k15_bvfunc_1(X0,X1,k10_eqrel_1(X0)),k15_bvfunc_1(X0,X1,X2)) & r1_tarski(k15_bvfunc_1(X0,X1,X2),k15_bvfunc_1(X0,X1,k6_partit1(X0))))))) [negated conjecture 1]
24. ! [X0] : m1_eqrel_1(k10_eqrel_1(X0),X0) [input dt_k10_eqrel_1]
26. ! [X0] : (~v1_xboole_0(X0) => m1_eqrel_1(k6_partit1(X0),X0)) [input dt_k6_partit1]
38. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_eqrel_1(X2,X0) => ! [X3] : (m1_eqrel_1(X3,X0) => (r1_setfam_1(X2,X3) => r1_tarski(k15_bvfunc_1(X0,X1,X2),k15_bvfunc_1(X0,X1,X3))))))) [input t1_bvfunc11]
39. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (m1_eqrel_1(X1,X0) => (r1_setfam_1(k10_eqrel_1(X0),X1) & r1_setfam_1(X1,k6_partit1(X0))))) [input t32_partit1]
55. ? [X0] : (? [X1] : (? [X2] : ((~r1_tarski(k15_bvfunc_1(X0,X1,k10_eqrel_1(X0)),k15_bvfunc_1(X0,X1,X2)) | ~r1_tarski(k15_bvfunc_1(X0,X1,X2),k15_bvfunc_1(X0,X1,k6_partit1(X0)))) & m1_eqrel_1(X2,X0)) & m1_subset_1(X1,X0)) & ~v1_xboole_0(X0)) [ennf transformation 2]
56. ! [X0] : (m1_eqrel_1(k6_partit1(X0),X0) | v1_xboole_0(X0)) [ennf transformation 26]
58. ! [X0] : (! [X1] : ((r1_setfam_1(k10_eqrel_1(X0),X1) & r1_setfam_1(X1,k6_partit1(X0))) | ~m1_eqrel_1(X1,X0)) | v1_xboole_0(X0)) [ennf transformation 39]
59. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_tarski(k15_bvfunc_1(X0,X1,X2),k15_bvfunc_1(X0,X1,X3)) | ~r1_setfam_1(X2,X3)) | ~m1_eqrel_1(X3,X0)) | ~m1_eqrel_1(X2,X0)) | ~m1_subset_1(X1,X0)) | v1_xboole_0(X0)) [ennf transformation 38]
60. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (r1_tarski(k15_bvfunc_1(X0,X1,X2),k15_bvfunc_1(X0,X1,X3)) | ~r1_setfam_1(X2,X3) | ~m1_eqrel_1(X3,X0)) | ~m1_eqrel_1(X2,X0)) | ~m1_subset_1(X1,X0)) | v1_xboole_0(X0)) [flattening 59]
86. ? [X0] : (? [X1] : (? [X2] : ((~r1_tarski(k15_bvfunc_1(X0,X1,k10_eqrel_1(X0)),k15_bvfunc_1(X0,X1,X2)) | ~r1_tarski(k15_bvfunc_1(X0,X1,X2),k15_bvfunc_1(X0,X1,k6_partit1(X0)))) & m1_eqrel_1(X2,X0)) & m1_subset_1(X1,X0)) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : ((~r1_tarski(k15_bvfunc_1(sK0,X1,k10_eqrel_1(sK0)),k15_bvfunc_1(sK0,X1,X2)) | ~r1_tarski(k15_bvfunc_1(sK0,X1,X2),k15_bvfunc_1(sK0,X1,k6_partit1(sK0)))) & m1_eqrel_1(X2,sK0)) & m1_subset_1(X1,sK0)) & ~v1_xboole_0(sK0)) [choice axiom]
87. ? [X1] : (? [X2] : ((~r1_tarski(k15_bvfunc_1(sK0,X1,k10_eqrel_1(sK0)),k15_bvfunc_1(sK0,X1,X2)) | ~r1_tarski(k15_bvfunc_1(sK0,X1,X2),k15_bvfunc_1(sK0,X1,k6_partit1(sK0)))) & m1_eqrel_1(X2,sK0)) & m1_subset_1(X1,sK0)) => (? [X2] : ((~r1_tarski(k15_bvfunc_1(sK0,sK1,k10_eqrel_1(sK0)),k15_bvfunc_1(sK0,sK1,X2)) | ~r1_tarski(k15_bvfunc_1(sK0,sK1,X2),k15_bvfunc_1(sK0,sK1,k6_partit1(sK0)))) & m1_eqrel_1(X2,sK0)) & m1_subset_1(sK1,sK0)) [choice axiom]
88. ? [X2] : ((~r1_tarski(k15_bvfunc_1(sK0,sK1,k10_eqrel_1(sK0)),k15_bvfunc_1(sK0,sK1,X2)) | ~r1_tarski(k15_bvfunc_1(sK0,sK1,X2),k15_bvfunc_1(sK0,sK1,k6_partit1(sK0)))) & m1_eqrel_1(X2,sK0)) => ((~r1_tarski(k15_bvfunc_1(sK0,sK1,k10_eqrel_1(sK0)),k15_bvfunc_1(sK0,sK1,sK2)) | ~r1_tarski(k15_bvfunc_1(sK0,sK1,sK2),k15_bvfunc_1(sK0,sK1,k6_partit1(sK0)))) & m1_eqrel_1(sK2,sK0)) [choice axiom]
89. (((~r1_tarski(k15_bvfunc_1(sK0,sK1,k10_eqrel_1(sK0)),k15_bvfunc_1(sK0,sK1,sK2)) | ~r1_tarski(k15_bvfunc_1(sK0,sK1,sK2),k15_bvfunc_1(sK0,sK1,k6_partit1(sK0)))) & m1_eqrel_1(sK2,sK0)) & m1_subset_1(sK1,sK0)) & ~v1_xboole_0(sK0) [skolemisation 55,88,87,86]
106. ~v1_xboole_0(sK0) [cnf transformation 89] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
107. m1_subset_1(sK1,sK0) [cnf transformation 89] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
108. m1_eqrel_1(sK2,sK0) [cnf transformation 89] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
109. ~r1_tarski(k15_bvfunc_1(sK0,sK1,k10_eqrel_1(sK0)),k15_bvfunc_1(sK0,sK1,sK2)) | ~r1_tarski(k15_bvfunc_1(sK0,sK1,sK2),k15_bvfunc_1(sK0,sK1,k6_partit1(sK0))) [cnf transformation 89] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
113. m1_eqrel_1(k10_eqrel_1(X0),X0) [cnf transformation 24] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
114. m1_eqrel_1(k6_partit1(X0),X0) | v1_xboole_0(X0) [cnf transformation 56] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
116. r1_setfam_1(X1,k6_partit1(X0)) | ~m1_eqrel_1(X1,X0) | v1_xboole_0(X0) [cnf transformation 58] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
117. r1_setfam_1(k10_eqrel_1(X0),X1) | ~m1_eqrel_1(X1,X0) | v1_xboole_0(X0) [cnf transformation 58] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
118. r1_tarski(k15_bvfunc_1(X0,X1,X2),k15_bvfunc_1(X0,X1,X3)) | ~r1_setfam_1(X2,X3) | ~m1_eqrel_1(X3,X0) | ~m1_eqrel_1(X2,X0) | ~m1_subset_1(X1,X0) | v1_xboole_0(X0) [cnf transformation 60] {a:0,w:23,wCS:9200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
148. 1 <=> r1_tarski(k15_bvfunc_1(sK0,sK1,sK2),k15_bvfunc_1(sK0,sK1,k6_partit1(sK0))) [avatar definition]
150. ~r1_tarski(k15_bvfunc_1(sK0,sK1,sK2),k15_bvfunc_1(sK0,sK1,k6_partit1(sK0))) <- (~1) [avatar component clause 148] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
152. 2 <=> r1_tarski(k15_bvfunc_1(sK0,sK1,k10_eqrel_1(sK0)),k15_bvfunc_1(sK0,sK1,sK2)) [avatar definition]
154. ~r1_tarski(k15_bvfunc_1(sK0,sK1,k10_eqrel_1(sK0)),k15_bvfunc_1(sK0,sK1,sK2)) <- (~2) [avatar component clause 152] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
155. ~1 | ~2 [avatar split clause 109,152,148]
157. 3 <=> m1_eqrel_1(sK2,sK0) [avatar definition]
159. m1_eqrel_1(sK2,sK0) <- (3) [avatar component clause 157] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
160. 3 [avatar split clause 108,157]
162. 4 <=> m1_subset_1(sK1,sK0) [avatar definition]
164. m1_subset_1(sK1,sK0) <- (4) [avatar component clause 162] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
165. 4 [avatar split clause 107,162]
167. 5 <=> v1_xboole_0(sK0) [avatar definition]
169. ~v1_xboole_0(sK0) <- (~5) [avatar component clause 167] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
170. ~5 [avatar split clause 106,167]
263. ~r1_setfam_1(sK2,k6_partit1(sK0)) | ~m1_eqrel_1(k6_partit1(sK0),sK0) | ~m1_eqrel_1(sK2,sK0) | ~m1_subset_1(sK1,sK0) | v1_xboole_0(sK0) <- (~1) [resolution 118,150] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
264. ~r1_setfam_1(sK2,k6_partit1(sK0)) | ~m1_eqrel_1(k6_partit1(sK0),sK0) | ~m1_subset_1(sK1,sK0) | v1_xboole_0(sK0) <- (~1, 3) [subsumption resolution 263,159] {a:1,w:13,wCS:1300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
265. ~r1_setfam_1(sK2,k6_partit1(sK0)) | ~m1_eqrel_1(k6_partit1(sK0),sK0) | v1_xboole_0(sK0) <- (~1, 3, 4) [subsumption resolution 264,164] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
266. ~r1_setfam_1(sK2,k6_partit1(sK0)) | ~m1_eqrel_1(k6_partit1(sK0),sK0) <- (~1, 3, 4, ~5) [subsumption resolution 265,169] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
268. 14 <=> m1_eqrel_1(k6_partit1(sK0),sK0) [avatar definition]
270. ~m1_eqrel_1(k6_partit1(sK0),sK0) <- (~14) [avatar component clause 268] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
272. 15 <=> r1_setfam_1(sK2,k6_partit1(sK0)) [avatar definition]
274. ~r1_setfam_1(sK2,k6_partit1(sK0)) <- (~15) [avatar component clause 272] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
275. ~14 | ~15 | 1 | ~3 | ~4 | 5 [avatar split clause 266,167,162,157,148,272,268]
277. v1_xboole_0(sK0) <- (~14) [resolution 270,114] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
279. $false <- (~5, ~14) [subsumption resolution 277,169] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
280. 5 | 14 [avatar contradiction clause 279]
283. ~m1_eqrel_1(sK2,sK0) | v1_xboole_0(sK0) <- (~15) [resolution 274,116] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
285. v1_xboole_0(sK0) <- (3, ~15) [subsumption resolution 283,159] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
286. $false <- (3, ~5, ~15) [subsumption resolution 285,169] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
287. ~3 | 5 | 15 [avatar contradiction clause 286]
288. ~r1_setfam_1(k10_eqrel_1(sK0),sK2) | ~m1_eqrel_1(sK2,sK0) | ~m1_eqrel_1(k10_eqrel_1(sK0),sK0) | ~m1_subset_1(sK1,sK0) | v1_xboole_0(sK0) <- (~2) [resolution 154,118] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
289. ~r1_setfam_1(k10_eqrel_1(sK0),sK2) | ~m1_eqrel_1(k10_eqrel_1(sK0),sK0) | ~m1_subset_1(sK1,sK0) | v1_xboole_0(sK0) <- (~2, 3) [subsumption resolution 288,159] {a:1,w:13,wCS:1300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
290. ~r1_setfam_1(k10_eqrel_1(sK0),sK2) | ~m1_subset_1(sK1,sK0) | v1_xboole_0(sK0) <- (~2, 3) [subsumption resolution 289,113] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
291. ~r1_setfam_1(k10_eqrel_1(sK0),sK2) | v1_xboole_0(sK0) <- (~2, 3, 4) [subsumption resolution 290,164] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
292. ~r1_setfam_1(k10_eqrel_1(sK0),sK2) <- (~2, 3, 4, ~5) [subsumption resolution 291,169] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
294. 16 <=> r1_setfam_1(k10_eqrel_1(sK0),sK2) [avatar definition]
296. ~r1_setfam_1(k10_eqrel_1(sK0),sK2) <- (~16) [avatar component clause 294] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
297. ~16 | 2 | ~3 | ~4 | 5 [avatar split clause 292,167,162,157,152,294]
299. ~m1_eqrel_1(sK2,sK0) | v1_xboole_0(sK0) <- (~16) [resolution 296,117] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
301. v1_xboole_0(sK0) <- (3, ~16) [subsumption resolution 299,159] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
302. $false <- (3, ~5, ~16) [subsumption resolution 301,169] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
303. ~3 | 5 | 16 [avatar contradiction clause 302]
304. $false [avatar sat refutation 155,160,165,170,275,280,287,297,303] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:134,thDist:-134}
% SZS output end Proof for bvfunc11__t4_bvfunc11
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 133
% Active clauses: 58
% Passive clauses: 77
% Final active clauses: 54
% Final passive clauses: 17
% Main loop iterations started: 60
% 
% >>> Simplifying Inferences
% Duplicate literals: 18
% Fw subsumption resolutions: 14
% Fw demodulations: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 41
% Unit resulting resolution: 3
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 22
% SAT solver unit clauses: 8
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward demodulation index maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

