i: [1,51,0,6,2,1,0,0] t4_wellord1
a: [3,92,0,2,1,51]
s: 92
i: [1,49,0,2,1,1,0,0] t4_wellord1
a: [3,99,0,2,1,49]
s: 99
a: [3,89,0,4,1,51]
s: 89
i: [1,83,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 83
i: [1,63,0,6,3,0,0,1] d4_wellord1
s: 63
i: [1,64,0,6,3,0,0,1] d4_wellord1
s: 64
r: 89
r: 92
a: [3,93,0,2,1,51]
s: 93
a: [3,88,0,4,1,51]
s: 88
d: [2,101,1,4,2,1,40,63,92]
d: [2,102,1,2,1,2,21,101,99]
a: [3,106,1,2,1,102]
s: 106
i: [1,65,0,6,3,0,0,1] d4_wellord1
s: 65
i: [1,66,0,6,3,0,0,1] d4_wellord1
s: 66
i: [1,67,0,6,3,0,0,1] d4_wellord1
s: 67
i: [1,78,0,6,3,0,0,1] cc1_wellord1
s: 78
i: [1,79,0,6,3,0,0,1] cc1_wellord1
s: 79
i: [1,80,0,6,3,0,0,1] cc1_wellord1
s: 80
i: [1,81,0,6,3,0,0,1] cc1_wellord1
s: 81
i: [1,82,0,6,3,0,0,1] cc1_wellord1
s: 82
i: [1,84,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 84
i: [1,53,0,8,3,0,0,2] d9_relat_2
s: 53
i: [1,54,0,8,3,0,0,2] d9_relat_2
s: 54
i: [1,55,0,8,3,0,0,2] d12_relat_2
s: 55
i: [1,56,0,8,3,0,0,2] d12_relat_2
s: 56
i: [1,57,0,8,3,0,0,2] d14_relat_2
s: 57
i: [1,58,0,8,3,0,0,2] d14_relat_2
s: 58
i: [1,59,0,8,3,0,0,2] d16_relat_2
s: 59
i: [1,60,0,8,3,0,0,2] d16_relat_2
s: 60
i: [1,61,0,8,3,0,0,2] t3_wellord1
s: 61
i: [1,62,0,8,3,0,0,2] t3_wellord1
s: 62
i: [1,69,0,8,3,0,0,1] d5_wellord1
s: 69
i: [1,70,0,8,3,0,0,1] d5_wellord1
s: 70
i: [1,71,0,8,3,0,0,1] d5_wellord1
s: 71
r: 88
r: 93
s: 92
r: 92
s: 93
d: [2,108,1,4,2,1,40,64,92]
d: [2,109,1,2,1,2,21,108,99]
a: [3,113,1,2,1,109]
s: 113
s: 88
i: [1,72,0,8,3,0,0,1] d5_wellord1
s: 72
r: 88
r: 93
s: 92
s: 89
d: [2,146,1,4,2,1,40,92,67]
d: [2,151,1,2,1,2,21,146,99]
a: [3,155,1,2,1,151]
s: 155
i: [1,73,0,8,3,0,0,1] d5_wellord1
s: 73
i: [1,85,0,8,3,0,0,255] fc3_relat_1
s: 85
d: [2,147,1,4,2,1,40,92,66]
d: [2,157,1,2,1,2,21,147,99]
a: [3,161,1,2,1,157]
s: 161
i: [1,52,0,10,2,0,0,255] d6_relat_1
s: 52
d: [2,148,1,4,2,1,40,92,65]
d: [2,163,1,2,1,2,21,148,99]
a: [3,167,1,2,1,163]
s: 167
i: [1,68,0,14,7,0,0,1] d4_wellord1
s: 68
i: [1,76,0,14,7,0,0,1] cc2_wellord1
s: 76
d: [2,125,1,6,2,1,40,69,88]
d: [2,126,1,4,1,2,21,125,99]
a: [3,130,1,4,1,126]
s: 130
i: [1,74,0,20,7,0,0,1] d5_wellord1
s: 74
d: [2,132,1,6,2,1,40,70,88]
d: [2,133,1,4,1,2,21,132,99]
a: [3,137,1,4,1,133]
s: 137
d: [2,139,1,6,2,1,40,71,88]
d: [2,140,1,4,1,2,21,139,99]
a: [3,144,1,4,1,140]
s: 144
d: [2,202,1,11,4,0,44,85,52]
s: 202
d: [2,205,1,26,8,0,40,74,61]
d: [2,206,1,24,7,0,18,205]
s: 206
d: [2,180,1,6,2,1,40,72,88]
d: [2,181,1,4,1,2,21,180,99]
a: [3,184,1,4,1,181]
s: 184
d: [2,218,2,0,0,3,52,99,161,184,57]
e: 218
d: [2,219,2,4,2,1,40,184,57]
d: [2,221,2,2,1,2,21,219,161]
d: [2,222,2,0,0,3,21,221,99]
e: 222
r: 130
r: 184
a: [3,129,1,4,1,126]
s: 129
d: [2,224,2,0,0,3,52,99,106,129,53]
e: 224
d: [2,225,2,4,2,1,40,129,53]
d: [2,227,2,2,1,2,21,225,106]
d: [2,228,2,0,0,3,21,227,99]
e: 228
r: 129
r: 144
a: [3,143,1,4,1,140]
s: 143
d: [2,230,2,0,0,3,52,99,167,143,55]
e: 230
d: [2,231,2,4,2,1,40,143,55]
d: [2,233,2,2,1,2,21,231,167]
d: [2,234,2,0,0,3,21,233,99]
e: 234
r: 137
r: 143
a: [3,136,1,4,1,133]
s: 136
d: [2,236,2,0,0,3,52,99,113,136,59]
e: 236
d: [2,237,2,4,2,1,40,136,59]
d: [2,239,2,2,1,2,21,237,113]
d: [2,240,2,0,0,3,21,239,99]
e: 240
r: 89
r: 92
r: 136
r: 161
s: 93
s: 88
a: [3,160,1,2,1,157]
s: 160
s: 130
s: 137
s: 144
a: [3,185,1,4,1,181]
s: 185
d: [2,276,2,0,0,3,52,99,160,185,58]
e: 276
d: [2,277,2,4,2,1,40,185,58]
d: [2,279,2,2,1,2,21,277,160]
d: [2,280,2,0,0,3,21,279,99]
e: 280
r: 113
r: 160
r: 155
s: 113
a: [3,154,1,2,1,151]
s: 154
s: 161
d: [2,253,1,6,2,1,40,88,73]
d: [2,258,1,4,1,2,21,253,99]
a: [3,262,1,4,1,258]
s: 262
d: [2,290,2,0,0,3,52,99,154,262,62]
e: 290
d: [2,291,2,4,2,1,40,262,62]
d: [2,294,2,2,1,2,21,291,154]
d: [2,295,2,0,0,3,21,294,99]
e: 295
r: 154
r: 167
r: 106
i: [1,50,0,6,2,1,0,0] t4_wellord1
d: [2,149,1,4,2,1,40,92,64]
d: [2,169,1,2,1,2,21,149,99]
d: [2,187,1,4,2,1,40,92,67]
d: [2,192,1,2,1,2,21,187,99]
d: [2,188,1,4,2,1,40,92,66]
d: [2,194,1,2,1,2,21,188,99]
d: [2,189,1,4,2,1,40,92,65]
d: [2,196,1,2,1,2,21,189,99]
d: [2,210,2,20,6,1,40,206,130]
d: [2,212,2,16,5,2,21,210,144]
d: [2,213,2,12,4,3,21,212,137]
d: [2,214,2,8,3,4,21,213,89]
d: [2,215,2,6,2,5,21,214,99]
d: [2,216,2,4,1,6,21,215,155]
d: [2,254,1,6,2,1,40,88,72]
d: [2,264,1,4,1,2,21,254,99]
d: [2,255,1,6,2,1,40,88,71]
d: [2,266,1,4,1,2,21,255,99]
d: [2,256,1,6,2,1,40,88,70]
d: [2,268,1,4,1,2,21,256,99]
d: [2,257,1,6,2,1,40,88,69]
d: [2,270,1,4,1,2,21,257,99]
d: [2,274,2,4,2,1,40,137,60]
d: [2,282,2,2,1,2,21,274,99]
d: [2,203,2,12,6,1,40,68,155]
d: [2,248,2,10,5,2,21,203,167]
d: [2,284,2,8,4,3,21,248,93]
d: [2,285,2,6,3,4,21,284,106]
d: [2,286,2,4,2,5,21,285,99]
d: [2,288,2,2,1,2,21,277,99]
d: [2,297,2,2,1,2,21,291,99]
d: [2,275,2,4,2,1,40,144,56]
d: [2,299,2,2,1,2,21,275,99]
d: [2,273,2,4,2,1,40,130,54]
d: [2,301,2,2,1,2,21,273,99]
f: 301,299,297,288,286,282,270,268,266,264,258,240,234,228,222,216,196,194,192,169,102,49,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for wellord1__t4_wellord1
% SZS output start Proof for wellord1__t4_wellord1
1. ! [X0] : (v1_relat_1(X0) => (r2_wellord1(X0,k1_relat_1(X0)) <=> v2_wellord1(X0))) [input t4_wellord1]
2. ~! [X0] : (v1_relat_1(X0) => (r2_wellord1(X0,k1_relat_1(X0)) <=> v2_wellord1(X0))) [negated conjecture 1]
6. ! [X0] : (v1_relat_1(X0) => (v4_relat_2(X0) <=> r4_relat_2(X0,k1_relat_1(X0)))) [input d12_relat_2]
7. ! [X0] : (v1_relat_1(X0) => (v6_relat_2(X0) <=> r6_relat_2(X0,k1_relat_1(X0)))) [input d14_relat_2]
8. ! [X0] : (v1_relat_1(X0) => (v8_relat_2(X0) <=> r8_relat_2(X0,k1_relat_1(X0)))) [input d16_relat_2]
9. ! [X0] : (v1_relat_1(X0) => (v2_wellord1(X0) <=> (v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d4_wellord1]
10. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r2_wellord1(X0,X1) <=> (r1_wellord1(X0,X1) & r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)))) [input d5_wellord1]
12. ! [X0] : (v1_relat_1(X0) => (v1_relat_2(X0) <=> r1_relat_2(X0,k1_relat_1(X0)))) [input d9_relat_2]
19. ! [X0] : (v1_relat_1(X0) => (v1_wellord1(X0) <=> r1_wellord1(X0,k1_relat_1(X0)))) [input t3_wellord1]
21. ? [X0] : ((r2_wellord1(X0,k1_relat_1(X0)) <~> v2_wellord1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
23. ! [X0] : ((v1_relat_2(X0) <=> r1_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 12]
24. ! [X0] : ((v4_relat_2(X0) <=> r4_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 6]
25. ! [X0] : ((v6_relat_2(X0) <=> r6_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 7]
26. ! [X0] : ((v8_relat_2(X0) <=> r8_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 8]
27. ! [X0] : ((v1_wellord1(X0) <=> r1_wellord1(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 19]
28. ! [X0] : ((v2_wellord1(X0) <=> (v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 9]
29. ! [X0] : (! [X1] : (r2_wellord1(X0,X1) <=> (r1_wellord1(X0,X1) & r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 10]
36. ? [X0] : (((~v2_wellord1(X0) | ~r2_wellord1(X0,k1_relat_1(X0))) & (v2_wellord1(X0) | r2_wellord1(X0,k1_relat_1(X0)))) & v1_relat_1(X0)) [nnf transformation 21]
37. ? [X0] : ((~v2_wellord1(X0) | ~r2_wellord1(X0,k1_relat_1(X0))) & (v2_wellord1(X0) | r2_wellord1(X0,k1_relat_1(X0))) & v1_relat_1(X0)) [flattening 36]
38. ? [X0] : ((~v2_wellord1(X0) | ~r2_wellord1(X0,k1_relat_1(X0))) & (v2_wellord1(X0) | r2_wellord1(X0,k1_relat_1(X0))) & v1_relat_1(X0)) => ((~v2_wellord1(sK0) | ~r2_wellord1(sK0,k1_relat_1(sK0))) & (v2_wellord1(sK0) | r2_wellord1(sK0,k1_relat_1(sK0))) & v1_relat_1(sK0)) [choice axiom]
39. (~v2_wellord1(sK0) | ~r2_wellord1(sK0,k1_relat_1(sK0))) & (v2_wellord1(sK0) | r2_wellord1(sK0,k1_relat_1(sK0))) & v1_relat_1(sK0) [skolemisation 37,38]
40. ! [X0] : (((v1_relat_2(X0) | ~r1_relat_2(X0,k1_relat_1(X0))) & (r1_relat_2(X0,k1_relat_1(X0)) | ~v1_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 23]
41. ! [X0] : (((v4_relat_2(X0) | ~r4_relat_2(X0,k1_relat_1(X0))) & (r4_relat_2(X0,k1_relat_1(X0)) | ~v4_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 24]
42. ! [X0] : (((v6_relat_2(X0) | ~r6_relat_2(X0,k1_relat_1(X0))) & (r6_relat_2(X0,k1_relat_1(X0)) | ~v6_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 25]
43. ! [X0] : (((v8_relat_2(X0) | ~r8_relat_2(X0,k1_relat_1(X0))) & (r8_relat_2(X0,k1_relat_1(X0)) | ~v8_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 26]
44. ! [X0] : (((v1_wellord1(X0) | ~r1_wellord1(X0,k1_relat_1(X0))) & (r1_wellord1(X0,k1_relat_1(X0)) | ~v1_wellord1(X0))) | ~v1_relat_1(X0)) [nnf transformation 27]
45. ! [X0] : (((v2_wellord1(X0) | (~v1_wellord1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_wellord1(X0))) | ~v1_relat_1(X0)) [nnf transformation 28]
46. ! [X0] : (((v2_wellord1(X0) | ~v1_wellord1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_wellord1(X0))) | ~v1_relat_1(X0)) [flattening 45]
47. ! [X0] : (! [X1] : ((r2_wellord1(X0,X1) | (~r1_wellord1(X0,X1) | ~r6_relat_2(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1))) & ((r1_wellord1(X0,X1) & r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r2_wellord1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 29]
48. ! [X0] : (! [X1] : ((r2_wellord1(X0,X1) | ~r1_wellord1(X0,X1) | ~r6_relat_2(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1)) & ((r1_wellord1(X0,X1) & r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r2_wellord1(X0,X1))) | ~v1_relat_1(X0)) [flattening 47]
49. v1_relat_1(sK0) [cnf transformation 39] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. v2_wellord1(sK0) | r2_wellord1(sK0,k1_relat_1(sK0)) [cnf transformation 39] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. ~v2_wellord1(sK0) | ~r2_wellord1(sK0,k1_relat_1(sK0)) [cnf transformation 39] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. r1_relat_2(X0,k1_relat_1(X0)) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
54. ~r1_relat_2(X0,k1_relat_1(X0)) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
55. r4_relat_2(X0,k1_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 41] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
56. ~r4_relat_2(X0,k1_relat_1(X0)) | v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 41] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
57. r6_relat_2(X0,k1_relat_1(X0)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 42] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
58. ~r6_relat_2(X0,k1_relat_1(X0)) | v6_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 42] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
59. r8_relat_2(X0,k1_relat_1(X0)) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 43] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
60. ~r8_relat_2(X0,k1_relat_1(X0)) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 43] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
61. r1_wellord1(X0,k1_relat_1(X0)) | ~v1_wellord1(X0) | ~v1_relat_1(X0) [cnf transformation 44] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
62. ~r1_wellord1(X0,k1_relat_1(X0)) | v1_wellord1(X0) | ~v1_relat_1(X0) [cnf transformation 44] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
63. ~v2_wellord1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. ~v2_wellord1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. ~v2_wellord1(X0) | v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
66. ~v2_wellord1(X0) | v6_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. ~v2_wellord1(X0) | v1_wellord1(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. ~v1_wellord1(X0) | v2_wellord1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. ~r2_wellord1(X0,X1) | r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 48] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. ~r2_wellord1(X0,X1) | r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 48] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. ~r2_wellord1(X0,X1) | r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 48] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. ~r2_wellord1(X0,X1) | r6_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 48] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. ~r2_wellord1(X0,X1) | r1_wellord1(X0,X1) | ~v1_relat_1(X0) [cnf transformation 48] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
74. ~r1_wellord1(X0,X1) | r2_wellord1(X0,X1) | ~r6_relat_2(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 48] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. 1 <=> r2_wellord1(sK0,k1_relat_1(sK0)) [avatar definition]
88. r2_wellord1(sK0,k1_relat_1(sK0)) <- (1) [avatar component clause 87] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
89. ~r2_wellord1(sK0,k1_relat_1(sK0)) <- (~1) [avatar component clause 87] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
91. 2 <=> v2_wellord1(sK0) [avatar definition]
92. v2_wellord1(sK0) <- (2) [avatar component clause 91] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
93. ~v2_wellord1(sK0) <- (~2) [avatar component clause 91] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
94. ~1 | ~2 [avatar split clause 51,91,87]
95. 1 | 2 [avatar split clause 50,91,87]
97. 3 <=> v1_relat_1(sK0) [avatar definition]
99. v1_relat_1(sK0) <- (3) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
100. 3 [avatar split clause 49,97]
101. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 63,92] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
102. v1_relat_2(sK0) <- (2, 3) [subsumption resolution 101,99] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
104. 4 <=> v1_relat_2(sK0) [avatar definition]
106. v1_relat_2(sK0) <- (4) [avatar component clause 104] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
107. 4 | ~2 | ~3 [avatar split clause 102,97,91,104]
111. 5 <=> v8_relat_2(sK0) [avatar definition]
113. v8_relat_2(sK0) <- (5) [avatar component clause 111] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. 6 <=> r1_relat_2(sK0,k1_relat_1(sK0)) [avatar definition]
129. ~r1_relat_2(sK0,k1_relat_1(sK0)) <- (~6) [avatar component clause 128] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
130. r1_relat_2(sK0,k1_relat_1(sK0)) <- (6) [avatar component clause 128] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
135. 7 <=> r8_relat_2(sK0,k1_relat_1(sK0)) [avatar definition]
136. ~r8_relat_2(sK0,k1_relat_1(sK0)) <- (~7) [avatar component clause 135] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
137. r8_relat_2(sK0,k1_relat_1(sK0)) <- (7) [avatar component clause 135] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. 8 <=> r4_relat_2(sK0,k1_relat_1(sK0)) [avatar definition]
143. ~r4_relat_2(sK0,k1_relat_1(sK0)) <- (~8) [avatar component clause 142] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. r4_relat_2(sK0,k1_relat_1(sK0)) <- (8) [avatar component clause 142] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
149. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 92,64] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
153. 9 <=> v1_wellord1(sK0) [avatar definition]
155. v1_wellord1(sK0) <- (9) [avatar component clause 153] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
159. 10 <=> v6_relat_2(sK0) [avatar definition]
161. v6_relat_2(sK0) <- (10) [avatar component clause 159] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
165. 11 <=> v4_relat_2(sK0) [avatar definition]
167. v4_relat_2(sK0) <- (11) [avatar component clause 165] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
169. v8_relat_2(sK0) <- (2, 3) [subsumption resolution 149,99] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
170. 5 | ~2 | ~3 [avatar split clause 169,97,91,111]
183. 12 <=> r6_relat_2(sK0,k1_relat_1(sK0)) [avatar definition]
184. ~r6_relat_2(sK0,k1_relat_1(sK0)) <- (~12) [avatar component clause 183] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
185. r6_relat_2(sK0,k1_relat_1(sK0)) <- (12) [avatar component clause 183] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
187. v1_wellord1(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 92,67] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
188. v6_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 92,66] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
189. v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 92,65] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
192. v1_wellord1(sK0) <- (2, 3) [subsumption resolution 187,99] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
193. 9 | ~2 | ~3 [avatar split clause 192,97,91,153]
194. v6_relat_2(sK0) <- (2, 3) [subsumption resolution 188,99] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
195. 10 | ~2 | ~3 [avatar split clause 194,97,91,159]
196. v4_relat_2(sK0) <- (2, 3) [subsumption resolution 189,99] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
197. 11 | ~2 | ~3 [avatar split clause 196,97,91,165]
203. v2_wellord1(sK0) | ~v6_relat_2(sK0) | ~v4_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (9) [resolution 68,155] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
205. r2_wellord1(X0,k1_relat_1(X0)) | ~r6_relat_2(X0,k1_relat_1(X0)) | ~r4_relat_2(X0,k1_relat_1(X0)) | ~r8_relat_2(X0,k1_relat_1(X0)) | ~r1_relat_2(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_wellord1(X0) | ~v1_relat_1(X0) [resolution 74,61] {a:1,w:26,wCS:10400,sine:1,thAx:0,allAx:2,thDist:-2}
206. ~r1_relat_2(X0,k1_relat_1(X0)) | ~r6_relat_2(X0,k1_relat_1(X0)) | ~r4_relat_2(X0,k1_relat_1(X0)) | ~r8_relat_2(X0,k1_relat_1(X0)) | r2_wellord1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_wellord1(X0) [duplicate literal removal 205] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
210. ~r6_relat_2(sK0,k1_relat_1(sK0)) | ~r4_relat_2(sK0,k1_relat_1(sK0)) | ~r8_relat_2(sK0,k1_relat_1(sK0)) | r2_wellord1(sK0,k1_relat_1(sK0)) | ~v1_relat_1(sK0) | ~v1_wellord1(sK0) <- (6) [resolution 206,130] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
212. ~r6_relat_2(sK0,k1_relat_1(sK0)) | ~r8_relat_2(sK0,k1_relat_1(sK0)) | r2_wellord1(sK0,k1_relat_1(sK0)) | ~v1_relat_1(sK0) | ~v1_wellord1(sK0) <- (6, 8) [subsumption resolution 210,144] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
213. ~r6_relat_2(sK0,k1_relat_1(sK0)) | r2_wellord1(sK0,k1_relat_1(sK0)) | ~v1_relat_1(sK0) | ~v1_wellord1(sK0) <- (6, 7, 8) [subsumption resolution 212,137] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
214. ~r6_relat_2(sK0,k1_relat_1(sK0)) | ~v1_relat_1(sK0) | ~v1_wellord1(sK0) <- (~1, 6, 7, 8) [subsumption resolution 213,89] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
215. ~r6_relat_2(sK0,k1_relat_1(sK0)) | ~v1_wellord1(sK0) <- (~1, 3, 6, 7, 8) [subsumption resolution 214,99] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
216. ~r6_relat_2(sK0,k1_relat_1(sK0)) <- (~1, 3, 6, 7, 8, 9) [subsumption resolution 215,155] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
217. ~12 | 1 | ~3 | ~6 | ~7 | ~8 | ~9 [avatar split clause 216,153,142,135,128,97,87,183]
219. ~v6_relat_2(sK0) | ~v1_relat_1(sK0) <- (~12) [resolution 184,57] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
221. ~v1_relat_1(sK0) <- (10, ~12) [subsumption resolution 219,161] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
222. $false <- (3, 10, ~12) [subsumption resolution 221,99] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
223. ~3 | ~10 | 12 [avatar contradiction clause 222]
225. ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (~6) [resolution 129,53] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
227. ~v1_relat_1(sK0) <- (4, ~6) [subsumption resolution 225,106] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
228. $false <- (3, 4, ~6) [subsumption resolution 227,99] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
229. ~3 | ~4 | 6 [avatar contradiction clause 228]
231. ~v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (~8) [resolution 143,55] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
233. ~v1_relat_1(sK0) <- (~8, 11) [subsumption resolution 231,167] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
234. $false <- (3, ~8, 11) [subsumption resolution 233,99] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
235. ~3 | 8 | ~11 [avatar contradiction clause 234]
237. ~v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (~7) [resolution 136,59] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
239. ~v1_relat_1(sK0) <- (5, ~7) [subsumption resolution 237,113] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
240. $false <- (3, 5, ~7) [subsumption resolution 239,99] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
241. ~3 | ~5 | 7 [avatar contradiction clause 240]
248. v2_wellord1(sK0) | ~v6_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (9, 11) [subsumption resolution 203,167] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
253. r1_wellord1(sK0,k1_relat_1(sK0)) | ~v1_relat_1(sK0) <- (1) [resolution 88,73] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
254. r6_relat_2(sK0,k1_relat_1(sK0)) | ~v1_relat_1(sK0) <- (1) [resolution 88,72] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
255. r4_relat_2(sK0,k1_relat_1(sK0)) | ~v1_relat_1(sK0) <- (1) [resolution 88,71] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
256. r8_relat_2(sK0,k1_relat_1(sK0)) | ~v1_relat_1(sK0) <- (1) [resolution 88,70] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
257. r1_relat_2(sK0,k1_relat_1(sK0)) | ~v1_relat_1(sK0) <- (1) [resolution 88,69] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
258. r1_wellord1(sK0,k1_relat_1(sK0)) <- (1, 3) [subsumption resolution 253,99] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
260. 13 <=> r1_wellord1(sK0,k1_relat_1(sK0)) [avatar definition]
262. r1_wellord1(sK0,k1_relat_1(sK0)) <- (13) [avatar component clause 260] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
263. 13 | ~1 | ~3 [avatar split clause 258,97,87,260]
264. r6_relat_2(sK0,k1_relat_1(sK0)) <- (1, 3) [subsumption resolution 254,99] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
265. 12 | ~1 | ~3 [avatar split clause 264,97,87,183]
266. r4_relat_2(sK0,k1_relat_1(sK0)) <- (1, 3) [subsumption resolution 255,99] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
267. 8 | ~1 | ~3 [avatar split clause 266,97,87,142]
268. r8_relat_2(sK0,k1_relat_1(sK0)) <- (1, 3) [subsumption resolution 256,99] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
269. 7 | ~1 | ~3 [avatar split clause 268,97,87,135]
270. r1_relat_2(sK0,k1_relat_1(sK0)) <- (1, 3) [subsumption resolution 257,99] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
271. 6 | ~1 | ~3 [avatar split clause 270,97,87,128]
273. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (6) [resolution 130,54] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
274. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (7) [resolution 137,60] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
275. v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (8) [resolution 144,56] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
277. v6_relat_2(sK0) | ~v1_relat_1(sK0) <- (12) [resolution 185,58] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
282. v8_relat_2(sK0) <- (3, 7) [subsumption resolution 274,99] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
283. 5 | ~3 | ~7 [avatar split clause 282,135,97,111]
284. ~v6_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (~2, 9, 11) [subsumption resolution 248,93] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
285. ~v6_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (~2, 4, 9, 11) [subsumption resolution 284,106] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
286. ~v6_relat_2(sK0) | ~v8_relat_2(sK0) <- (~2, 3, 4, 9, 11) [subsumption resolution 285,99] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
287. ~5 | ~10 | 2 | ~3 | ~4 | ~9 | ~11 [avatar split clause 286,165,153,104,97,91,159,111]
288. v6_relat_2(sK0) <- (3, 12) [subsumption resolution 277,99] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
289. 10 | ~3 | ~12 [avatar split clause 288,183,97,159]
291. v1_wellord1(sK0) | ~v1_relat_1(sK0) <- (13) [resolution 262,62] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
297. v1_wellord1(sK0) <- (3, 13) [subsumption resolution 291,99] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
298. 9 | ~3 | ~13 [avatar split clause 297,260,97,153]
299. v4_relat_2(sK0) <- (3, 8) [subsumption resolution 275,99] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
300. 11 | ~3 | ~8 [avatar split clause 299,142,97,165]
301. v1_relat_2(sK0) <- (3, 6) [subsumption resolution 273,99] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
302. 4 | ~3 | ~6 [avatar split clause 301,128,97,104]
303. $false [avatar sat refutation 94,95,100,107,170,193,195,197,217,223,229,235,241,263,265,267,269,271,283,287,289,298,300,302] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:295,thDist:-295}
% SZS output end Proof for wellord1__t4_wellord1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 1
% 
% >>> Saturation
% Initial clauses: 37
% Generated clauses: 204
% Active clauses: 61
% Passive clauses: 74
% Final active clauses: 38
% Final passive clauses: 4
% Main loop iterations started: 67
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Fw subsumption resolutions: 55
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Forward subsumptions: 20
% 
% >>> Generating Inferences
% Binary resolution: 47
% Unit resulting resolution: 6
% Forward superposition: 3
% Backward superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 47
% SAT solver unit clauses: 1
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4989
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% splitting model update: 0.001 s ( own 0.0000 s ) 
% binary resolution index maintenance: 0.001 s
% other: 0.004 s ( own 0.003 s ) 

