i: [1,21,0,3,1,1,0,0] t101_zfmisc_1
a: [3,36,0,3,1,21]
s: 36
i: [1,20,0,3,1,1,0,0] t101_zfmisc_1
a: [3,41,0,3,1,20]
s: 41
i: [1,22,0,11,1,1,0,0] t101_zfmisc_1
a: [3,31,0,11,1,22]
s: 31
i: [1,25,0,3,1,0,0,1] reflexivity_r1_tarski
s: 25
i: [1,26,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 26
i: [1,27,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 27
i: [1,23,0,8,2,0,0,1] t28_xboole_1
s: 23
i: [1,24,0,15,1,0,0,1] t100_zfmisc_1
s: 24
d: [2,45,1,14,2,1,44,31,23]
a: [3,51,1,7,1,45]
s: 51
d: [2,61,1,11,1,1,44,31,24]
d: [2,67,1,11,1,1,22,61,27]
a: [3,71,1,11,1,67]
s: 71
d: [2,43,1,8,2,0,44,23,27]
s: 43
d: [2,95,2,14,2,1,44,31,43]
a: [3,102,2,7,1,95]
s: 102
d: [2,44,1,8,2,0,44,23,27]
s: 44
d: [2,74,2,12,2,1,44,71,23]
d: [2,76,2,9,1,2,21,74,41]
a: [3,80,2,9,1,76]
s: 80
d: [2,127,3,10,2,1,44,80,23]
d: [2,128,3,3,1,1,19,127]
d: [2,130,3,0,0,2,21,128,36]
e: 130
f: 130,76,67,20,21,22
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t101_zfmisc_1
% SZS output start Proof for zfmisc_1__t101_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => k3_xboole_0(k2_zfmisc_1(X0,X3),k2_zfmisc_1(X1,X2)) = k2_zfmisc_1(X0,X2)) [input t101_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => k3_xboole_0(k2_zfmisc_1(X0,X3),k2_zfmisc_1(X1,X2)) = k2_zfmisc_1(X0,X2)) [negated conjecture 1]
3. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
8. ! [X0] : ! [X1] : ! [X2] : ! [X3] : k2_zfmisc_1(k3_xboole_0(X0,X1),k3_xboole_0(X2,X3)) = k3_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) [input t100_zfmisc_1]
9. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
10. ~! [X0,X1,X2,X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => k3_xboole_0(k2_zfmisc_1(X0,X3),k2_zfmisc_1(X1,X2)) = k2_zfmisc_1(X0,X2)) [flattening 2]
11. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 9]
12. ! [X0,X1,X2,X3] : k2_zfmisc_1(k3_xboole_0(X0,X1),k3_xboole_0(X2,X3)) = k3_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) [flattening 8]
15. ? [X0,X1,X2,X3] : (k3_xboole_0(k2_zfmisc_1(X0,X3),k2_zfmisc_1(X1,X2)) != k2_zfmisc_1(X0,X2) & (r1_tarski(X2,X3) & r1_tarski(X0,X1))) [ennf transformation 10]
16. ? [X0,X1,X2,X3] : (k3_xboole_0(k2_zfmisc_1(X0,X3),k2_zfmisc_1(X1,X2)) != k2_zfmisc_1(X0,X2) & r1_tarski(X2,X3) & r1_tarski(X0,X1)) [flattening 15]
17. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 11]
18. ? [X0,X1,X2,X3] : (k3_xboole_0(k2_zfmisc_1(X0,X3),k2_zfmisc_1(X1,X2)) != k2_zfmisc_1(X0,X2) & r1_tarski(X2,X3) & r1_tarski(X0,X1)) => (k3_xboole_0(k2_zfmisc_1(sK0,sK3),k2_zfmisc_1(sK1,sK2)) != k2_zfmisc_1(sK0,sK2) & r1_tarski(sK2,sK3) & r1_tarski(sK0,sK1)) [choice axiom]
19. k3_xboole_0(k2_zfmisc_1(sK0,sK3),k2_zfmisc_1(sK1,sK2)) != k2_zfmisc_1(sK0,sK2) & r1_tarski(sK2,sK3) & r1_tarski(sK0,sK1) [skolemisation 16,18]
20. r1_tarski(sK0,sK1) [cnf transformation 19] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
21. r1_tarski(sK2,sK3) [cnf transformation 19] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
22. k3_xboole_0(k2_zfmisc_1(sK0,sK3),k2_zfmisc_1(sK1,sK2)) != k2_zfmisc_1(sK0,sK2) [cnf transformation 19] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
23. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 17] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
24. k2_zfmisc_1(k3_xboole_0(X0,X1),k3_xboole_0(X2,X3)) = k3_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) [cnf transformation 12] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
27. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
29. 1 <=> k3_xboole_0(k2_zfmisc_1(sK0,sK3),k2_zfmisc_1(sK1,sK2)) = k2_zfmisc_1(sK0,sK2) [avatar definition]
31. k3_xboole_0(k2_zfmisc_1(sK0,sK3),k2_zfmisc_1(sK1,sK2)) != k2_zfmisc_1(sK0,sK2) <- (~1) [avatar component clause 29] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
32. ~1 [avatar split clause 22,29]
34. 2 <=> r1_tarski(sK2,sK3) [avatar definition]
36. r1_tarski(sK2,sK3) <- (2) [avatar component clause 34] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
37. 2 [avatar split clause 21,34]
39. 3 <=> r1_tarski(sK0,sK1) [avatar definition]
41. r1_tarski(sK0,sK1) <- (3) [avatar component clause 39] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
42. 3 [avatar split clause 20,39]
61. k2_zfmisc_1(sK0,sK2) != k2_zfmisc_1(k3_xboole_0(sK0,sK1),k3_xboole_0(sK3,sK2)) <- (~1) [superposition 31,24] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
67. k2_zfmisc_1(sK0,sK2) != k2_zfmisc_1(k3_xboole_0(sK0,sK1),k3_xboole_0(sK2,sK3)) <- (~1) [forward demodulation 61,27] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
69. 6 <=> k2_zfmisc_1(sK0,sK2) = k2_zfmisc_1(k3_xboole_0(sK0,sK1),k3_xboole_0(sK2,sK3)) [avatar definition]
71. k2_zfmisc_1(sK0,sK2) != k2_zfmisc_1(k3_xboole_0(sK0,sK1),k3_xboole_0(sK2,sK3)) <- (~6) [avatar component clause 69] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. ~6 | 1 [avatar split clause 67,29,69]
74. k2_zfmisc_1(sK0,sK2) != k2_zfmisc_1(sK0,k3_xboole_0(sK2,sK3)) | ~r1_tarski(sK0,sK1) <- (~6) [superposition 71,23] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
76. k2_zfmisc_1(sK0,sK2) != k2_zfmisc_1(sK0,k3_xboole_0(sK2,sK3)) <- (3, ~6) [subsumption resolution 74,41] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
78. 7 <=> k2_zfmisc_1(sK0,sK2) = k2_zfmisc_1(sK0,k3_xboole_0(sK2,sK3)) [avatar definition]
80. k2_zfmisc_1(sK0,sK2) != k2_zfmisc_1(sK0,k3_xboole_0(sK2,sK3)) <- (~7) [avatar component clause 78] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
81. ~7 | ~3 | 6 [avatar split clause 76,69,39,78]
127. k2_zfmisc_1(sK0,sK2) != k2_zfmisc_1(sK0,sK2) | ~r1_tarski(sK2,sK3) <- (~7) [superposition 80,23] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
128. ~r1_tarski(sK2,sK3) <- (~7) [trivial inequality removal 127] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
130. $false <- (2, ~7) [subsumption resolution 128,36] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
131. ~2 | 7 [avatar contradiction clause 130]
132. $false [avatar sat refutation 32,37,42,72,81,131] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:71,thDist:-71}
% SZS output end Proof for zfmisc_1__t101_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 8
% Generated clauses: 54
% Active clauses: 13
% Passive clauses: 29
% Final active clauses: 13
% Final passive clauses: 15
% Main loop iterations started: 14
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 3
% Fw demodulations: 3
% 
% >>> Deletion Inferences
% Forward subsumptions: 5
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Forward superposition: 14
% Backward superposition: 15
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 3
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

