i: [1,441,0,2,1,1,0,0] t1_complex1
a: [3,706,0,2,1,441]
s: 706
i: [1,440,0,2,1,1,0,0] t1_complex1
a: [3,711,0,2,1,440]
s: 711
i: [1,443,0,3,1,1,0,0] t1_complex1
a: [3,696,0,3,1,443]
s: 696
i: [1,442,0,7,1,1,0,0] t1_complex1
a: [3,701,0,7,1,442]
s: 701
i: [1,444,0,2,1,0,0,2] fc1_numbers
a: [3,716,0,2,1,444]
s: 716
i: [1,445,0,2,1,0,0,3] spc2_boole
a: [3,721,0,2,1,445]
s: 721
i: [1,446,0,2,1,0,0,3] spc1_boole
a: [3,726,0,2,1,446]
s: 726
i: [1,447,0,2,1,0,0,255] fc3_membered
a: [3,731,0,2,1,447]
s: 731
i: [1,448,0,2,1,0,0,255] fc6_membered
a: [3,736,0,2,1,448]
s: 736
i: [1,449,0,2,1,0,0,255] fc1_xboole_0
a: [3,741,0,2,1,449]
s: 741
i: [1,450,0,2,1,0,0,3] spc0_boole
a: [3,746,0,2,1,450]
s: 746
i: [1,535,0,2,1,0,0,2] spc2_numerals
a: [3,1176,0,2,1,535]
s: 1176
i: [1,539,0,2,1,0,0,2] spc1_numerals
a: [3,1196,0,2,1,539]
s: 1196
i: [1,662,0,2,1,0,0,3] rc2_xboole_0
a: [3,1223,0,2,1,662]
s: 1223
i: [1,663,0,2,1,0,0,1] rc1_xreal_0
a: [3,1228,0,2,1,663]
s: 1228
i: [1,664,0,2,1,0,0,255] rc1_xxreal_0
a: [3,1233,0,2,1,664]
s: 1233
i: [1,665,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,1238,0,2,1,665]
s: 1238
i: [1,666,0,2,1,0,0,3] rc1_xboole_0
a: [3,1243,0,2,1,666]
s: 1243
i: [1,668,0,2,1,0,0,3] rc1_membered
a: [3,1248,0,2,1,668]
s: 1248
i: [1,667,0,2,1,0,0,3] rc1_membered
a: [3,1253,0,2,1,667]
s: 1253
i: [1,670,0,2,1,0,0,3] rc2_membered
a: [3,1258,0,2,1,670]
s: 1258
i: [1,669,0,2,1,0,0,3] rc2_membered
a: [3,1263,0,2,1,669]
s: 1263
i: [1,672,0,2,1,0,0,255] rc3_membered
a: [3,1268,0,2,1,672]
s: 1268
i: [1,671,0,2,1,0,0,255] rc3_membered
a: [3,1273,0,2,1,671]
s: 1273
i: [1,674,0,2,1,0,0,3] rc2_xcmplx_0
a: [3,1278,0,2,1,674]
s: 1278
i: [1,673,0,2,1,0,0,3] rc2_xcmplx_0
a: [3,1283,0,2,1,673]
s: 1283
i: [1,676,0,2,1,0,0,255] rc3_xxreal_0
a: [3,1288,0,2,1,676]
s: 1288
i: [1,675,0,2,1,0,0,255] rc3_xxreal_0
a: [3,1293,0,2,1,675]
s: 1293
i: [1,678,0,2,1,0,0,255] rc2_xxreal_0
a: [3,1298,0,2,1,678]
s: 1298
i: [1,677,0,2,1,0,0,255] rc2_xxreal_0
a: [3,1303,0,2,1,677]
s: 1303
i: [1,682,0,2,1,0,0,255] rc3_xreal_0
a: [3,1308,0,2,1,682]
s: 1308
i: [1,681,0,2,1,0,0,255] rc3_xreal_0
a: [3,1313,0,2,1,681]
s: 1313
i: [1,680,0,2,1,0,0,255] rc3_xreal_0
a: [3,1318,0,2,1,680]
s: 1318
i: [1,679,0,2,1,0,0,255] rc3_xreal_0
a: [3,1323,0,2,1,679]
s: 1323
i: [1,686,0,2,1,0,0,255] rc2_xreal_0
a: [3,1328,0,2,1,686]
s: 1328
i: [1,685,0,2,1,0,0,255] rc2_xreal_0
a: [3,1333,0,2,1,685]
s: 1333
i: [1,684,0,2,1,0,0,255] rc2_xreal_0
a: [3,1338,0,2,1,684]
s: 1338
i: [1,683,0,2,1,0,0,255] rc2_xreal_0
a: [3,1343,0,2,1,683]
s: 1343
i: [1,688,0,2,1,0,0,255] rc4_xxreal_0
a: [3,1348,0,2,1,688]
s: 1348
i: [1,687,0,2,1,0,0,255] rc4_xxreal_0
a: [3,1353,0,2,1,687]
s: 1353
i: [1,692,0,2,1,0,0,255] rc4_xreal_0
a: [3,1358,0,2,1,692]
s: 1358
i: [1,691,0,2,1,0,0,255] rc4_xreal_0
a: [3,1363,0,2,1,691]
s: 1363
i: [1,690,0,2,1,0,0,255] rc4_xreal_0
a: [3,1368,0,2,1,690]
s: 1368
i: [1,689,0,2,1,0,0,255] rc4_xreal_0
a: [3,1373,0,2,1,689]
s: 1373
i: [1,623,0,3,1,0,0,255] reflexivity_r1_tarski
s: 623
i: [1,451,0,3,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r2_r0
a: [3,751,0,3,1,451]
s: 751
i: [1,452,0,3,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r2_r1
a: [3,756,0,3,1,452]
s: 756
i: [1,453,0,3,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r1_r0
a: [3,761,0,3,1,453]
s: 761
i: [1,454,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,766,0,3,1,454]
s: 766
r: 741
d: [2,1375,0,2,1,2,23,741,766]
a: [3,1379,0,2,1,1375]
s: 1379
i: [1,455,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,771,0,3,1,455]
s: 771
r: 736
d: [2,1381,0,2,1,2,23,736,771]
a: [3,1385,0,2,1,1381]
s: 1385
i: [1,456,0,3,1,0,0,255] t1_numerals
a: [3,776,0,3,1,456]
s: 776
i: [1,457,0,3,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r0_r0
a: [3,781,0,3,1,457]
s: 781
i: [1,458,0,3,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r0_r2
a: [3,786,0,3,1,458]
s: 786
i: [1,459,0,3,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r0_r1
a: [3,791,0,3,1,459]
s: 791
i: [1,460,0,3,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r2_r2
a: [3,796,0,3,1,460]
s: 796
i: [1,461,0,3,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r1_r2
a: [3,801,0,3,1,461]
s: 801
i: [1,462,0,3,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,806,0,3,1,462]
s: 806
i: [1,538,0,3,1,0,0,2] spc2_numerals
a: [3,1161,0,3,1,538]
s: 1161
i: [1,537,0,3,1,0,0,2] spc2_numerals
a: [3,1166,0,3,1,537]
s: 1166
i: [1,542,0,3,1,0,0,2] spc1_numerals
a: [3,1181,0,3,1,542]
s: 1181
i: [1,541,0,3,1,0,0,2] spc1_numerals
a: [3,1186,0,3,1,541]
s: 1186
i: [1,545,0,3,1,0,0,2] spc0_numerals
a: [3,1201,0,3,1,545]
s: 1201
i: [1,544,0,3,1,0,0,2] spc0_numerals
a: [3,1206,0,3,1,544]
s: 1206
d: [2,1387,0,3,1,2,22,776,771]
d: [2,1388,0,3,1,3,22,1387,766]
a: [3,1392,0,3,1,1388]
s: 1392
i: [1,546,0,4,2,0,0,255] cc2_xreal_0
s: 546
i: [1,547,0,4,2,0,0,255] cc2_xxreal_0
s: 547
i: [1,548,0,4,2,0,0,255] cc2_xcmplx_0
s: 548
i: [1,553,0,4,2,0,0,255] cc2_membered
s: 553
i: [1,554,0,4,2,0,0,255] cc3_membered
s: 554
i: [1,557,0,4,2,0,0,255] cc5_membered
s: 557
i: [1,558,0,4,2,0,0,255] cc4_membered
s: 558
d: [2,1394,1,2,1,1,40,557,731]
a: [3,1398,1,2,1,1394]
s: 1398
i: [1,561,0,4,2,0,0,255] cc1_membered
s: 561
i: [1,564,0,4,2,0,0,255] cc4_xreal_0
s: 564
d: [2,1430,1,2,1,1,40,564,711]
a: [3,1439,1,2,1,1430]
s: 1439
i: [1,565,0,4,2,0,0,2] cc3_xreal_0
s: 565
d: [2,1431,1,2,1,1,40,564,706]
a: [3,1444,1,2,1,1431]
s: 1444
i: [1,587,0,4,2,0,0,3] cc18_membered
s: 587
i: [1,622,0,4,1,0,0,3] existence_m1_subset_1
s: 622
d: [2,1451,1,2,1,1,40,565,711]
a: [3,1460,1,2,1,1451]
s: 1460
i: [1,463,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r0_rm2
a: [3,811,0,4,1,463]
s: 811
d: [2,1452,1,2,1,1,40,565,706]
a: [3,1465,1,2,1,1452]
s: 1465
i: [1,464,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r0_rm1
a: [3,816,0,4,1,464]
s: 816
i: [1,465,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r2_rm2
a: [3,821,0,4,1,465]
s: 821
d: [2,1400,1,2,1,1,40,558,731]
a: [3,1404,1,2,1,1400]
s: 1404
i: [1,466,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r2_rm1
a: [3,826,0,4,1,466]
s: 826
d: [2,1406,1,2,1,1,40,561,1385]
a: [3,1413,1,2,1,1406]
s: 1413
i: [1,467,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r1_rm2
a: [3,831,0,4,1,467]
s: 831
i: [1,468,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__r1_rm1
a: [3,836,0,4,1,468]
s: 836
d: [2,1407,1,2,1,1,40,561,1248]
a: [3,1418,1,2,1,1407]
s: 1418
i: [1,469,0,4,1,0,0,2] rqRealNeg__k4_xcmplx_0__r0_r0
a: [3,841,0,4,1,469]
s: 841
d: [2,1408,1,2,1,1,40,561,1258]
a: [3,1423,1,2,1,1408]
s: 1423
i: [1,470,0,4,1,0,0,2] dt_k5_numbers
a: [3,846,0,4,1,470]
s: 846
i: [1,471,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm2_r0
a: [3,851,0,4,1,471]
s: 851
d: [2,1409,1,2,1,1,40,561,1268]
a: [3,1428,1,2,1,1409]
s: 1428
i: [1,472,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm2_r2
a: [3,856,0,4,1,472]
s: 856
d: [2,1432,1,2,1,1,40,564,1228]
a: [3,1449,1,2,1,1432]
s: 1449
i: [1,473,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm2_r1
a: [3,861,0,4,1,473]
s: 861
i: [1,474,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm1_r0
a: [3,866,0,4,1,474]
s: 866
d: [2,1453,1,2,1,1,40,565,1228]
a: [3,1470,1,2,1,1453]
s: 1470
i: [1,475,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm1_r2
a: [3,871,0,4,1,475]
s: 871
d: [2,1472,1,2,1,1,40,587,1379]
a: [3,1480,1,2,1,1472]
s: 1480
i: [1,476,0,4,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm1_r1
a: [3,876,0,4,1,476]
s: 876
i: [1,477,0,4,1,0,0,1] dt_k6_numbers
a: [3,881,0,4,1,477]
s: 881
d: [2,1473,1,2,1,1,40,587,746]
a: [3,1485,1,2,1,1473]
s: 1485
i: [1,536,0,4,1,0,0,2] spc2_numerals
a: [3,1171,0,4,1,536]
s: 1171
d: [2,1474,1,2,1,1,40,587,1243]
a: [3,1490,1,2,1,1474]
s: 1490
i: [1,540,0,4,1,0,0,2] spc1_numerals
a: [3,1191,0,4,1,540]
s: 1191
i: [1,543,0,4,1,0,0,2] spc0_numerals
a: [3,1211,0,4,1,543]
s: 1211
d: [2,1475,1,2,1,1,40,587,1353]
a: [3,1495,1,2,1,1475]
s: 1495
i: [1,566,0,5,2,0,0,1] fc2_square_1
s: 566
d: [2,1476,1,2,1,1,40,587,1373]
a: [3,1500,1,2,1,1476]
s: 1500
i: [1,568,0,5,2,0,0,2] fc3_xreal_0
s: 568
i: [1,569,0,5,2,0,0,2] fc3_xreal_0
s: 569
d: [2,1502,2,2,1,1,40,1413,553]
a: [3,1506,2,2,1,1502]
s: 1506
i: [1,578,0,5,2,0,0,1] fc1_square_1
s: 578
d: [2,1508,2,2,1,1,40,1418,553]
a: [3,1512,2,2,1,1508]
s: 1512
i: [1,579,0,5,2,0,0,2] dt_k4_xcmplx_0
s: 579
i: [1,588,0,5,2,0,0,255] t6_boole
s: 588
d: [2,1514,2,2,1,1,40,1423,553]
a: [3,1518,2,2,1,1514]
s: 1518
i: [1,589,0,5,2,0,0,2] cc25_membered
s: 589
d: [2,1584,1,2,1,1,40,589,1186]
a: [3,1590,1,2,1,1584]
s: 1590
i: [1,590,0,5,2,0,0,2] cc1_xreal_0
s: 590
i: [1,591,0,5,2,0,0,2] cc1_xcmplx_0
s: 591
d: [2,1585,1,2,1,1,40,589,1166]
a: [3,1595,1,2,1,1585]
s: 1595
i: [1,619,0,5,2,0,0,255] t7_boole
s: 619
d: [2,1608,1,2,1,1,40,590,1201]
a: [3,1615,1,2,1,1608]
s: 1615
i: [1,478,0,5,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm1_rm2
a: [3,886,0,5,1,478]
s: 886
i: [1,479,0,5,1,0,0,2] rqRealDiff__k6_xcmplx_0__r0_r0_r0
a: [3,891,0,5,1,479]
s: 891
d: [2,1609,1,2,1,1,40,590,1181]
a: [3,1620,1,2,1,1609]
s: 1620
i: [1,480,0,5,1,0,0,2] rqRealDiff__k6_xcmplx_0__r2_r2_r0
a: [3,896,0,5,1,480]
s: 896
d: [2,1610,1,2,1,1,40,590,1161]
a: [3,1625,1,2,1,1610]
s: 1625
i: [1,481,0,5,1,0,0,2] rqRealDiff__k6_xcmplx_0__r1_r1_r0
a: [3,901,0,5,1,481]
s: 901
i: [1,482,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r0_r0_r0
a: [3,906,0,5,1,482]
s: 906
d: [2,1632,1,2,1,1,40,591,1201]
a: [3,1639,1,2,1,1632]
s: 1639
i: [1,483,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r0_r2_r0
a: [3,911,0,5,1,483]
s: 911
d: [2,1633,1,2,1,1,40,591,1181]
a: [3,1644,1,2,1,1633]
s: 1644
i: [1,484,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r0_r1_r0
a: [3,916,0,5,1,484]
s: 916
i: [1,485,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r2_r0_r0
a: [3,921,0,5,1,485]
s: 921
d: [2,1634,1,2,1,1,40,591,1161]
a: [3,1649,1,2,1,1634]
s: 1649
i: [1,486,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_r0_r0
a: [3,926,0,5,1,486]
s: 926
d: [2,1520,2,2,1,1,40,1428,553]
a: [3,1524,2,2,1,1520]
s: 1524
i: [1,487,0,5,1,0,0,1] rqRealAdd__k2_xcmplx_0__r0_r0_r0
a: [3,931,0,5,1,487]
s: 931
i: [1,488,0,5,1,0,0,2] rqRealNeg__k4_xcmplx_0__rm2_r2
a: [3,936,0,5,1,488]
s: 936
d: [2,1526,2,2,1,1,40,1480,561]
a: [3,1530,2,2,1,1526]
s: 1530
i: [1,489,0,5,1,0,0,2] rqRealDiff__k6_xcmplx_0__r2_r0_r2
a: [3,941,0,5,1,489]
s: 941
d: [2,1532,2,2,1,1,40,1485,561]
a: [3,1536,2,2,1,1532]
s: 1536
i: [1,490,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r2_r1_r2
a: [3,946,0,5,1,490]
s: 946
i: [1,491,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_r2_r2
a: [3,951,0,5,1,491]
s: 951
d: [2,1538,2,2,1,1,40,1490,561]
a: [3,1542,2,2,1,1538]
s: 1542
i: [1,492,0,5,1,0,0,1] rqRealAdd__k2_xcmplx_0__r0_r2_r2
a: [3,956,0,5,1,492]
s: 956
d: [2,1544,2,2,1,1,40,1495,561]
a: [3,1548,2,2,1,1544]
s: 1548
i: [1,493,0,5,1,0,0,1] rqRealAdd__k2_xcmplx_0__r2_r0_r2
a: [3,961,0,5,1,493]
s: 961
i: [1,494,0,5,1,0,0,1] rqRealAdd__k2_xcmplx_0__r1_r1_r2
a: [3,966,0,5,1,494]
s: 966
d: [2,1552,2,2,1,1,40,1500,561]
a: [3,1556,2,2,1,1552]
s: 1556
i: [1,495,0,5,1,0,0,2] rqRealNeg__k4_xcmplx_0__rm1_r1
a: [3,971,0,5,1,495]
s: 971
d: [2,1602,2,2,1,1,40,1590,561]
a: [3,1606,2,2,1,1602]
s: 1606
i: [1,496,0,5,1,0,0,2] rqRealDiff__k6_xcmplx_0__r2_r1_r1
a: [3,976,0,5,1,496]
s: 976
i: [1,497,0,5,1,0,0,2] rqRealDiff__k6_xcmplx_0__r1_r0_r1
a: [3,981,0,5,1,497]
s: 981
d: [2,1656,2,2,1,1,40,1595,561]
a: [3,1660,2,2,1,1656]
s: 1660
i: [1,498,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,986,0,5,1,498]
s: 986
d: [2,1663,2,2,1,1,40,1615,564]
a: [3,1668,2,2,1,1663]
s: 1668
i: [1,499,0,5,1,0,0,1] rqRealAdd__k2_xcmplx_0__r0_r1_r1
a: [3,991,0,5,1,499]
s: 991
i: [1,500,0,5,1,0,0,1] rqRealAdd__k2_xcmplx_0__r1_r0_r1
a: [3,996,0,5,1,500]
s: 996
d: [2,1671,2,2,1,1,40,1620,564]
a: [3,1676,2,2,1,1671]
s: 1676
i: [1,503,0,5,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm2_rm2
a: [3,1001,0,5,1,503]
s: 1001
d: [2,1679,2,2,1,1,40,1625,564]
a: [3,1684,2,2,1,1679]
s: 1684
i: [1,504,0,5,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm2_rm1
a: [3,1006,0,5,1,504]
s: 1006
i: [1,505,0,5,1,0,0,2] rqLessOrEqual__r1_xxreal_0__rm1_rm1
a: [3,1011,0,5,1,505]
s: 1011
d: [2,1562,3,2,1,1,40,1506,554]
a: [3,1566,3,2,1,1562]
s: 1566
i: [1,641,0,7,3,0,0,255] reflexivity_r1_xxreal_0
a: [3,1218,0,5,2,641]
s: 1218
d: [2,1568,3,2,1,1,40,1512,554]
a: [3,1572,3,2,1,1568]
s: 1572
d: [2,1575,0,5,2,1,22,588,766]
s: 1575
i: [1,592,0,6,2,0,0,2] cc11_membered
s: 592
d: [2,1576,3,2,1,1,40,1518,554]
a: [3,1580,3,2,1,1576]
s: 1580
i: [1,593,0,6,2,0,0,255] cc8_membered
s: 593
d: [2,1686,3,2,1,1,40,1524,554]
a: [3,1690,3,2,1,1686]
s: 1690
i: [1,606,0,6,3,0,0,255] cc5_xxreal_0
s: 606
i: [1,608,0,6,3,0,0,255] cc5_xxreal_0
s: 608
d: [2,1807,1,4,2,1,40,606,1288]
d: [2,1809,1,2,1,2,21,1807,1293]
a: [3,1813,1,2,1,1809]
s: 1813
i: [1,609,0,6,3,0,0,255] cc3_xxreal_0
s: 609
d: [2,1808,1,4,2,1,40,606,1313]
d: [2,1815,1,2,1,2,21,1808,1318]
a: [3,1819,1,2,1,1815]
s: 1819
i: [1,611,0,6,3,0,0,255] cc3_xxreal_0
s: 611
i: [1,613,0,6,3,0,0,255] cc7_xxreal_0
s: 613
d: [2,1821,1,4,2,1,40,608,1288]
d: [2,1823,1,2,1,2,21,1821,1293]
a: [3,1827,1,2,1,1823]
s: 1827
i: [1,614,0,6,3,0,0,255] cc7_xxreal_0
s: 614
d: [2,1822,1,4,2,1,40,608,1313]
d: [2,1829,1,2,1,2,21,1822,1318]
a: [3,1833,1,2,1,1829]
s: 1833
i: [1,615,0,6,2,0,0,255] t1_subset
s: 615
i: [1,624,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 624
d: [2,1837,1,4,2,1,40,609,1298]
d: [2,1839,1,2,1,2,21,1837,1303]
a: [3,1843,1,2,1,1839]
s: 1843
i: [1,506,0,6,1,0,0,2] rqRealMult__k3_xcmplx_0__r0_rm2_r0
a: [3,1016,0,6,1,506]
s: 1016
d: [2,1838,1,4,2,1,40,609,1333]
d: [2,1845,1,2,1,2,21,1838,1338]
a: [3,1849,1,2,1,1845]
s: 1849
i: [1,507,0,6,1,0,0,2] rqRealMult__k3_xcmplx_0__rm2_r0_r0
a: [3,1021,0,6,1,507]
s: 1021
i: [1,508,0,6,1,0,0,1] rqRealAdd__k2_xcmplx_0__r2_rm2_r0
a: [3,1026,0,6,1,508]
s: 1026
d: [2,1695,3,2,1,1,40,1530,553]
a: [3,1699,3,2,1,1695]
s: 1699
i: [1,509,0,6,1,0,0,1] rqRealAdd__k2_xcmplx_0__r1_rm1_r0
a: [3,1031,0,6,1,509]
s: 1031
d: [2,1701,3,2,1,1,40,1536,553]
a: [3,1705,3,2,1,1701]
s: 1705
i: [1,510,0,6,1,0,0,1] rqRealAdd__k2_xcmplx_0__rm2_r2_r0
a: [3,1036,0,6,1,510]
s: 1036
i: [1,511,0,6,1,0,0,1] rqRealAdd__k2_xcmplx_0__rm1_r1_r0
a: [3,1041,0,6,1,511]
s: 1041
d: [2,1707,3,2,1,1,40,1542,553]
a: [3,1711,3,2,1,1707]
s: 1711
i: [1,512,0,6,1,0,0,2] rqRealDiff__k6_xcmplx_0__r0_rm2_r2
a: [3,1046,0,6,1,512]
s: 1046
d: [2,1713,3,2,1,1,40,1548,553]
a: [3,1717,3,2,1,1713]
s: 1717
i: [1,513,0,6,1,0,0,2] rqRealDiff__k6_xcmplx_0__r1_rm1_r2
a: [3,1051,0,6,1,513]
s: 1051
i: [1,514,0,6,1,0,0,2] rqRealDiff__k6_xcmplx_0__r0_rm1_r1
a: [3,1056,0,6,1,514]
s: 1056
d: [2,1719,3,2,1,1,40,1556,553]
a: [3,1723,3,2,1,1719]
s: 1723
i: [1,515,0,6,1,0,0,1] rqRealAdd__k2_xcmplx_0__r2_rm1_r1
a: [3,1061,0,6,1,515]
s: 1061
d: [2,1728,3,2,1,1,40,1606,553]
a: [3,1732,3,2,1,1728]
s: 1732
i: [1,516,0,6,1,0,0,1] rqRealAdd__k2_xcmplx_0__rm1_r2_r1
a: [3,1066,0,6,1,516]
s: 1066
i: [1,517,0,6,1,0,0,2] rqRealDiff__k6_xcmplx_0__r0_r2_rm2
a: [3,1071,0,6,1,517]
s: 1071
d: [2,1734,3,2,1,1,40,1660,553]
a: [3,1738,3,2,1,1734]
s: 1738
i: [1,518,0,6,1,0,0,2] rqRealDiff__k6_xcmplx_0__r0_r1_rm1
a: [3,1076,0,6,1,518]
s: 1076
d: [2,1740,4,2,1,1,40,1566,558]
a: [3,1745,4,2,1,1740]
s: 1745
i: [1,519,0,6,1,0,0,2] rqRealDiff__k6_xcmplx_0__r1_r2_rm1
a: [3,1081,0,6,1,519]
s: 1081
i: [1,549,0,7,3,0,0,255] cc13_membered
s: 549
d: [2,1894,1,4,2,1,40,549,1392]
d: [2,1903,1,2,1,2,21,1894,1745]
a: [3,1907,1,2,1,1903]
s: 1907
i: [1,550,0,7,3,0,0,255] cc12_membered
s: 550
d: [2,1919,1,4,2,1,40,550,1392]
a: [3,1931,1,2,1,1919]
s: 1931
i: [1,559,0,7,3,0,0,255] cc14_membered
s: 559
i: [1,562,0,7,3,0,0,255] cc17_membered
s: 562
d: [2,1943,1,4,2,1,40,559,1392]
d: [2,1952,1,2,1,2,21,1943,1566]
a: [3,1956,1,2,1,1952]
s: 1956
i: [1,567,0,7,2,0,0,1] t63_xreal_1
s: 567
d: [2,1967,1,4,2,1,40,562,1392]
d: [2,1976,1,2,1,2,21,1967,1385]
a: [3,1980,1,2,1,1976]
s: 1980
i: [1,580,0,7,2,0,0,1] t2_arithm
s: 580
i: [1,581,0,7,2,0,0,2] involutiveness_k4_xcmplx_0
s: 581
d: [2,1969,1,4,2,1,40,562,1201]
a: [3,1998,1,2,1,1969]
s: 1998
i: [1,582,0,7,2,0,0,1] t4_arithm
s: 582
d: [2,1971,1,4,2,1,40,562,1181]
a: [3,2005,1,2,1,1971]
s: 2005
i: [1,583,0,7,2,0,0,2] t3_arithm
s: 583
i: [1,584,0,7,2,0,0,1] t1_arithm
s: 584
d: [2,1973,1,4,2,1,40,562,1161]
a: [3,2012,1,2,1,1973]
s: 2012
i: [1,598,0,7,3,0,0,3] fc6_xcmplx_0
s: 598
d: [2,1741,4,2,1,1,40,1566,557]
a: [3,1750,4,2,1,1741]
s: 1750
i: [1,599,0,7,3,0,0,3] fc6_xcmplx_0
s: 599
i: [1,600,0,7,3,0,0,3] fc16_xreal_0
s: 600
d: [2,1752,4,2,1,1,40,1572,558]
a: [3,1757,4,2,1,1752]
s: 1757
i: [1,601,0,7,3,0,0,3] fc16_xreal_0
s: 601
d: [2,2071,1,6,3,1,44,601,841]
d: [2,2075,1,4,2,2,21,2071,1615]
a: [3,2083,1,2,1,2075]
s: 2083
i: [1,602,0,7,3,0,0,3] fc15_xreal_0
s: 602
i: [1,603,0,7,3,0,0,3] fc15_xreal_0
s: 603
a: [3,2078,1,2,1,2075]
s: 2078
i: [1,617,0,7,2,0,0,255] t3_subset
s: 617
d: [2,1753,4,2,1,1,40,1572,557]
a: [3,1762,4,2,1,1753]
s: 1762
i: [1,618,0,7,3,0,0,3] t8_boole
s: 618
i: [1,520,0,7,1,0,0,2] rqRealDiff__k6_xcmplx_0__rm2_rm2_r0
a: [3,1086,0,7,1,520]
s: 1086
d: [2,1776,4,2,1,1,40,1580,558]
a: [3,1781,4,2,1,1776]
s: 1781
i: [1,521,0,7,1,0,0,2] rqRealDiff__k6_xcmplx_0__rm1_rm1_r0
a: [3,1091,0,7,1,521]
s: 1091
d: [2,1777,4,2,1,1,40,1580,557]
a: [3,1786,4,2,1,1777]
s: 1786
i: [1,522,0,7,1,0,0,2] rqRealDiff__k6_xcmplx_0__rm1_rm2_r1
a: [3,1096,0,7,1,522]
s: 1096
i: [1,523,0,7,1,0,0,2] rqRealDiff__k6_xcmplx_0__rm2_r0_rm2
a: [3,1101,0,7,1,523]
s: 1101
d: [2,1795,4,2,1,1,40,1690,558]
a: [3,1800,4,2,1,1795]
s: 1800
i: [1,524,0,7,1,0,0,2] rqRealDiff__k6_xcmplx_0__rm1_r1_rm2
a: [3,1106,0,7,1,524]
s: 1106
d: [2,1796,4,2,1,1,40,1690,557]
a: [3,1805,4,2,1,1796]
s: 1805
i: [1,525,0,7,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_rm2_rm2
a: [3,1111,0,7,1,525]
s: 1111
i: [1,526,0,7,1,0,0,2] rqRealMult__k3_xcmplx_0__rm2_r1_rm2
a: [3,1116,0,7,1,526]
s: 1116
d: [2,1851,4,2,1,1,40,1699,554]
a: [3,1856,4,2,1,1851]
s: 1856
i: [1,527,0,7,1,0,0,1] rqRealAdd__k2_xcmplx_0__r0_rm2_rm2
a: [3,1121,0,7,1,527]
s: 1121
d: [2,1858,4,2,1,1,40,1705,554]
a: [3,1862,4,2,1,1858]
s: 1862
i: [1,528,0,7,1,0,0,2] rqRealDiff__k6_xcmplx_0__rm1_r0_rm1
a: [3,1126,0,7,1,528]
s: 1126
i: [1,529,0,7,1,0,0,1] rqRealAdd__k2_xcmplx_0__r0_rm1_rm1
a: [3,1131,0,7,1,529]
s: 1131
d: [2,1864,4,2,1,1,40,1711,554]
a: [3,1868,4,2,1,1864]
s: 1868
i: [1,530,0,7,1,0,0,1] rqRealAdd__k2_xcmplx_0__r1_rm2_rm1
a: [3,1136,0,7,1,530]
s: 1136
d: [2,1870,4,2,1,1,40,1717,554]
a: [3,1874,4,2,1,1870]
s: 1874
i: [1,531,0,7,1,0,0,1] rqRealAdd__k2_xcmplx_0__rm2_r1_rm1
a: [3,1141,0,7,1,531]
s: 1141
i: [1,532,0,7,1,0,0,1] rqRealAdd__k2_xcmplx_0__rm1_r0_rm1
a: [3,1146,0,7,1,532]
s: 1146
d: [2,1876,4,2,1,1,40,1723,554]
a: [3,1880,4,2,1,1876]
s: 1880
i: [1,551,0,8,3,0,0,255] cc20_membered
s: 551
d: [2,1882,4,2,1,1,40,1732,554]
a: [3,1886,4,2,1,1882]
s: 1886
i: [1,552,0,8,3,0,0,255] cc19_membered
s: 552
i: [1,555,0,8,3,0,0,255] cc23_membered
s: 555
d: [2,1888,4,2,1,1,40,1738,554]
a: [3,1892,4,2,1,1888]
s: 1892
i: [1,556,0,8,3,0,0,255] cc22_membered
s: 556
d: [2,2136,5,2,1,1,40,1856,558]
a: [3,2142,5,2,1,2136]
s: 2142
i: [1,560,0,8,3,0,0,255] cc21_membered
s: 560
i: [1,563,0,8,3,0,0,3] cc24_membered
s: 563
d: [2,2137,5,2,1,1,40,1856,557]
a: [3,2147,5,2,1,2137]
s: 2147
i: [1,585,0,8,2,0,0,1] d1_square_1
s: 585
d: [2,2149,5,2,1,1,40,1862,558]
a: [3,2154,5,2,1,2149]
s: 2154
i: [1,595,0,8,4,0,0,255] cc4_xxreal_0
s: 595
i: [1,597,0,8,4,0,0,255] cc6_xxreal_0
s: 597
d: [2,2287,2,6,3,1,40,595,1439]
a: [3,2300,2,2,1,2287]
s: 2300
d: [2,2349,3,0,0,3,52,1379,696,2300,618]
e: 2349
d: [2,2350,3,0,0,3,52,1379,696,2300,618]
e: 2350
d: [2,2351,3,0,0,3,52,696,2300,1575]
e: 2351
r: 2300
i: [1,604,0,8,4,0,0,255] cc8_xxreal_0
s: 604
a: [3,2299,2,2,1,2287]
s: 2299
i: [1,616,0,8,3,0,0,255] t2_subset
s: 616
i: [1,638,0,8,3,0,0,2] fc7_xreal_0
s: 638
a: [3,2304,2,2,1,2287]
s: 2304
i: [1,639,0,8,3,0,0,2] fc6_xreal_0
s: 639
d: [2,2288,2,6,3,1,40,595,1444]
a: [3,2313,2,2,1,2288]
s: 2313
r: 2313
i: [1,640,0,8,3,0,0,1] fc5_xreal_0
s: 640
i: [1,643,0,8,3,0,0,2] fc4_xcmplx_0
s: 643
a: [3,2312,2,2,1,2288]
s: 2312
i: [1,644,0,8,3,0,0,2] fc3_xcmplx_0
s: 644
a: [3,2317,2,2,1,2288]
s: 2317
i: [1,645,0,8,3,0,0,1] fc2_xcmplx_0
s: 645
i: [1,533,0,8,1,0,0,1] rqRealAdd__k2_xcmplx_0__rm1_rm1_rm2
a: [3,1151,0,8,1,533]
s: 1151
d: [2,2290,1,6,3,1,40,595,1233]
a: [3,2339,1,2,1,2290]
s: 2339
r: 2339
i: [1,534,0,8,1,0,0,2] rqRealDiff__k6_xcmplx_0__rm2_rm1_rm1
a: [3,1156,0,8,1,534]
s: 1156
a: [3,2338,1,2,1,2290]
s: 2338
i: [1,586,0,9,2,0,0,2] spc2_arithm
s: 586
i: [1,621,0,9,3,0,0,255] t5_subset
s: 621
a: [3,2343,1,2,1,2290]
s: 2343
i: [1,620,0,10,3,0,0,255] t4_subset
s: 620
d: [2,2357,1,5,2,1,40,616,1392]
a: [3,2373,1,2,1,2357]
s: 2373
i: [1,642,0,10,4,0,0,255] connectedness_r1_xxreal_0
s: 642
i: [1,572,0,11,5,0,0,3] t2_real
s: 572
d: [2,2106,1,8,3,1,44,603,936]
a: [3,2119,1,2,1,2106]
s: 2119
i: [1,573,0,11,5,0,0,3] t1_real
s: 573
d: [2,2107,1,8,3,1,44,603,971]
a: [3,2128,1,2,1,2107]
s: 2128
i: [1,574,0,11,5,0,0,3] t8_real
s: 574
i: [1,575,0,11,5,0,0,3] t7_real
s: 575
d: [2,2628,1,14,5,0,40,572,567]
d: [2,2652,1,12,4,1,21,2628,1956]
d: [2,2653,1,8,3,1,21,2652,639]
a: [3,2657,1,2,1,2653]
s: 2657
d: [2,2746,2,0,0,3,52,1907,1379,2657,606]
e: 2746
d: [2,2748,2,4,2,1,40,2657,606]
d: [2,2752,2,2,1,2,21,2748,1379]
d: [2,2753,2,0,0,3,21,2752,1907]
e: 2753
r: 2657
i: [1,576,0,11,5,0,0,3] t3_real
s: 576
a: [3,2656,1,2,1,2653]
s: 2656
i: [1,577,0,11,5,0,0,3] t4_real
s: 577
i: [1,646,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 646
d: [2,2684,1,14,5,0,40,573,567]
d: [2,2708,1,12,4,1,21,2684,1956]
d: [2,2709,1,8,3,1,21,2708,639]
a: [3,2716,1,2,1,2709]
s: 2716
i: [1,647,0,11,3,0,0,1] commutativity_k2_xcmplx_0
s: 647
d: [2,2289,2,6,3,1,40,595,1449]
a: [3,2326,2,2,1,2289]
s: 2326
r: 2326
i: [1,629,0,12,5,0,0,2] fc8_xcmplx_0
s: 629
i: [1,630,0,12,5,0,0,2] fc26_xreal_0
s: 630
a: [3,2325,2,2,1,2289]
s: 2325
i: [1,631,0,12,5,0,0,3] fc9_xreal_0
s: 631
a: [3,2330,2,2,1,2289]
s: 2330
i: [1,632,0,12,5,0,0,2] fc17_xreal_0
s: 632
i: [1,633,0,12,5,0,0,2] fc18_xreal_0
s: 633
d: [2,2150,5,2,1,1,40,1862,557]
a: [3,2159,5,2,1,2150]
s: 2159
i: [1,634,0,12,5,0,0,2] fc24_xreal_0
s: 634
d: [2,2161,5,2,1,1,40,1868,558]
a: [3,2166,5,2,1,2161]
s: 2166
i: [1,635,0,12,5,0,0,2] fc23_xreal_0
s: 635
i: [1,636,0,12,5,0,0,2] fc25_xreal_0
s: 636
d: [2,2162,5,2,1,1,40,1868,557]
a: [3,2171,5,2,1,2162]
s: 2171
i: [1,637,0,12,5,0,0,3] fc10_xreal_0
s: 637
d: [2,2173,5,2,1,1,40,1874,558]
a: [3,2178,5,2,1,2173]
s: 2178
i: [1,648,0,12,3,0,0,2] spc1_arithm
s: 648
i: [1,651,0,12,5,0,0,2] fc21_xreal_0
s: 651
d: [2,2174,5,2,1,1,40,1874,557]
a: [3,2183,5,2,1,2174]
s: 2183
i: [1,652,0,12,5,0,0,2] fc22_xreal_0
s: 652
d: [2,2185,5,2,1,1,40,1880,558]
a: [3,2190,5,2,1,2185]
s: 2190
i: [1,653,0,12,5,0,0,3] fc13_xreal_0
s: 653
i: [1,654,0,12,5,0,0,3] fc14_xreal_0
s: 654
d: [2,3251,1,14,5,1,44,653,701]
d: [2,3279,1,12,4,2,21,3251,2656]
a: [3,3283,1,3,1,3279]
s: 3283
i: [1,655,0,12,5,0,0,3] fc12_xreal_0
s: 655
a: [3,3287,1,3,1,3279]
s: 3287
d: [2,3392,2,0,0,2,52,711,3287,566]
e: 3392
d: [2,3393,2,2,1,1,40,3287,566]
d: [2,3395,2,0,0,2,21,3393,711]
e: 3395
r: 3287
i: [1,656,0,12,5,0,0,3] fc11_xreal_0
s: 656
i: [1,657,0,12,5,0,0,2] fc20_xreal_0
s: 657
r: 2078
r: 2083
a: [3,3286,1,3,1,3279]
s: 3286
i: [1,658,0,12,5,0,0,2] fc19_xreal_0
s: 658
a: [3,3295,1,3,1,3279]
s: 3295
d: [2,3635,2,0,0,2,52,706,3295,566]
e: 3635
d: [2,3636,2,2,1,1,40,3295,566]
d: [2,3638,2,0,0,2,21,3636,706]
e: 3638
r: 3295
i: [1,570,0,13,6,0,0,3] t5_real
s: 570
i: [1,571,0,13,6,0,0,3] t6_real
s: 571
a: [3,3290,1,3,1,3279]
s: 3290
i: [1,649,0,13,3,0,0,2] spc9_arithm
s: 649
a: [3,3294,1,3,1,3279]
s: 3294
i: [1,628,0,14,4,0,0,2] existence_m2_subset_1
s: 628
i: [1,650,0,14,3,0,0,1] spc8_arithm
s: 650
d: [2,3305,1,14,5,1,44,654,701]
d: [2,3329,1,12,4,2,21,3305,2656]
a: [3,3333,1,3,1,3329]
s: 3333
i: [1,625,0,15,5,0,0,2] dt_m2_subset_1
s: 625
a: [3,3336,1,3,1,3329]
s: 3336
i: [1,626,0,15,5,0,0,2] redefinition_m2_subset_1
s: 626
i: [1,627,0,15,5,0,0,2] redefinition_m2_subset_1
s: 627
d: [2,3592,2,3,1,1,40,3286,565]
a: [3,3597,2,3,1,3592]
s: 3597
i: [1,659,0,17,4,0,0,2] spc7_arithm
s: 659
d: [2,3593,2,3,1,1,40,3286,564]
a: [3,3602,2,3,1,3593]
s: 3602
i: [1,660,0,17,4,0,0,1] spc6_arithm
s: 660
i: [1,661,0,19,4,0,0,2] spc5_arithm
s: 661
d: [2,3769,2,3,1,1,40,3294,565]
a: [3,3774,2,3,1,3769]
s: 3774
s: 2078
d: [2,3770,2,3,1,1,40,3294,564]
a: [3,3779,2,3,1,3770]
s: 3779
s: 2083
d: [2,1586,1,3,1,0,40,589,622]
a: [3,1600,1,3,1,1586]
s: 1600
d: [2,3978,3,9,3,1,40,3602,595]
d: [2,3979,3,6,2,2,21,3978,3283]
d: [2,3980,3,3,1,3,21,3979,3336]
a: [3,3984,3,3,1,3980]
s: 3984
d: [2,1611,1,3,1,0,40,590,622]
a: [3,1630,1,3,1,1611]
s: 1630
d: [2,4302,3,9,3,1,40,3779,595]
d: [2,4303,3,6,2,2,21,4302,3333]
d: [2,4304,3,3,1,3,21,4303,3290]
a: [3,4308,3,3,1,4304]
s: 4308
d: [2,1635,1,3,1,0,40,591,622]
a: [3,1654,1,3,1,1635]
s: 1654
d: [2,1895,1,5,2,1,40,549,846]
a: [3,1912,1,3,1,1895]
s: 1912
d: [2,4316,4,3,1,1,40,3984,587]
a: [3,4320,4,3,1,4316]
s: 4320
d: [2,1920,1,5,2,1,40,550,846]
a: [3,1936,1,3,1,1920]
s: 1936
d: [2,4330,4,3,1,1,40,4308,587]
a: [3,4334,4,3,1,4330]
s: 4334
d: [2,1944,1,5,2,1,40,559,846]
a: [3,1961,1,3,1,1944]
s: 1961
d: [2,1968,1,5,2,1,40,562,846]
a: [3,1985,1,3,1,1968]
s: 1985
d: [2,4336,5,3,1,1,40,4320,561]
a: [3,4340,5,3,1,4336]
s: 4340
d: [2,2020,1,5,2,1,44,567,906]
d: [2,2022,1,3,1,2,21,2020,1615]
a: [3,2026,1,3,1,2022]
s: 2026
d: [2,4342,5,3,1,1,40,4334,561]
a: [3,4346,5,3,1,4342]
s: 4346
d: [2,2021,1,5,2,1,44,567,986]
d: [2,2028,1,3,1,2,21,2021,1620]
a: [3,2032,1,3,1,2028]
s: 2032
d: [2,2039,1,7,3,0,44,567,580]
d: [2,2040,1,5,2,1,21,2039,1956]
d: [2,2041,1,3,1,2,21,2040,1931]
a: [3,2045,1,3,1,2041]
s: 2045
d: [2,4348,6,3,1,1,40,4340,553]
a: [3,4352,6,3,1,4348]
s: 4352
d: [2,2072,1,8,3,1,44,601,936]
d: [2,2085,1,6,2,2,21,2072,1176]
a: [3,2089,1,3,1,2085]
s: 2089
d: [2,4359,6,3,1,1,40,4346,553]
a: [3,4363,6,3,1,4359]
s: 4363
a: [3,2093,1,3,1,2085]
s: 2093
d: [2,4400,2,0,0,2,52,1625,2093,569]
e: 4400
d: [2,4401,2,2,1,1,40,2093,569]
d: [2,4403,2,0,0,2,21,4401,1625]
e: 4403
r: 2093
a: [3,2092,1,3,1,2085]
s: 2092
d: [2,4379,7,3,1,1,40,4352,554]
a: [3,4383,7,3,1,4379]
s: 4383
d: [2,2073,1,8,3,1,44,601,971]
d: [2,2095,1,6,2,2,21,2073,1196]
a: [3,2099,1,3,1,2095]
s: 2099
d: [2,4394,7,3,1,1,40,4363,554]
a: [3,4398,7,3,1,4394]
s: 4398
a: [3,2103,1,3,1,2095]
s: 2103
d: [2,4442,2,0,0,2,52,1620,2103,569]
e: 4442
d: [2,4443,2,2,1,1,40,2103,569]
d: [2,4445,2,0,0,2,21,4443,1620]
e: 4445
r: 2103
a: [3,2102,1,3,1,2095]
s: 2102
d: [2,4409,8,3,1,1,40,4383,558]
a: [3,4414,8,3,1,4409]
s: 4414
a: [3,2114,1,3,1,2106]
s: 2114
d: [2,4410,8,3,1,1,40,4383,557]
a: [3,4419,8,3,1,4410]
s: 4419
a: [3,2123,1,3,1,2107]
s: 2123
d: [2,2359,1,5,2,1,40,616,1201]
d: [2,2384,1,3,1,2,21,2359,716]
a: [3,2388,1,3,1,2384]
s: 2388
d: [2,4430,8,3,1,1,40,4398,558]
a: [3,4435,8,3,1,4430]
s: 4435
d: [2,2361,1,5,2,1,40,616,1181]
d: [2,2395,1,3,1,2,21,2361,716]
a: [3,2399,1,3,1,2395]
s: 2399
d: [2,4431,8,3,1,1,40,4398,557]
a: [3,4440,8,3,1,4431]
s: 4440
d: [2,2363,1,5,2,1,40,616,1161]
d: [2,2406,1,3,1,2,21,2363,716]
a: [3,2410,1,3,1,2406]
s: 2410
d: [2,2501,1,8,3,1,44,643,1106]
d: [2,2520,1,6,2,2,21,2501,1644]
a: [3,2523,1,3,1,2520]
s: 2523
d: [2,2186,5,2,1,1,40,1880,557]
a: [3,2195,5,2,1,2186]
s: 2195
a: [3,2528,1,3,1,2520]
s: 2528
d: [2,2200,5,2,1,1,40,1886,558]
a: [3,2205,5,2,1,2200]
s: 2205
d: [2,3644,1,15,5,1,40,570,1001]
d: [2,3663,1,12,4,1,18,3644]
a: [3,3678,1,3,1,3663]
s: 3678
d: [2,3645,1,15,5,1,40,570,1006]
a: [3,3683,1,3,1,3645]
s: 3683
d: [2,2201,5,2,1,1,40,1886,557]
a: [3,2210,5,2,1,2201]
s: 2210
d: [2,3990,1,19,4,1,44,660,701]
d: [2,4042,1,16,3,2,21,3990,3597]
a: [3,4045,1,13,2,4042]
s: 4045
d: [2,4475,2,14,3,1,44,4045,584]
d: [2,4490,2,12,2,2,21,4475,1931]
d: [2,4491,2,9,1,3,21,4490,3774]
d: [2,4492,2,5,1,4,22,4491,701]
a: [3,4496,2,5,1,4492]
s: 4496
d: [2,1770,1,4,1,0,40,592,622]
a: [3,1774,1,4,1,1770]
s: 1774
d: [2,1789,1,4,1,0,40,593,622]
a: [3,1793,1,4,1,1789]
s: 1793
d: [2,4513,3,7,2,2,44,4496,1575]
s: 4513
d: [2,2241,1,8,3,0,44,585,580]
d: [2,2253,1,6,2,0,18,2241]
d: [2,2254,1,4,1,1,21,2253,1931]
a: [3,2258,1,4,1,2254]
s: 2258
d: [2,4567,4,16,3,3,44,4045,4513]
d: [2,4589,4,13,2,4,21,4567,3774]
d: [2,4590,4,10,1,5,21,4589,4308]
d: [2,4591,4,6,1,6,22,4590,701]
a: [3,4595,4,6,1,4591]
s: 4595
d: [2,2242,1,6,2,1,44,585,906]
d: [2,2260,1,4,1,2,21,2242,1639]
a: [3,2264,1,4,1,2260]
s: 2264
d: [2,2243,1,6,2,1,44,585,986]
d: [2,2266,1,4,1,2,21,2243,1644]
a: [3,2270,1,4,1,2266]
s: 2270
d: [2,4601,5,8,2,2,44,4595,1575]
s: 4601
d: [2,2358,1,7,2,1,40,616,846]
a: [3,2378,1,4,1,2358]
s: 2378
d: [2,2218,5,2,1,1,40,1892,558]
a: [3,2223,5,2,1,2218]
s: 2223
a: [3,2381,1,3,1,2358]
s: 2381
d: [2,1550,1,5,2,0,40,566,565]
s: 1550
d: [2,2219,5,2,1,1,40,1892,557]
a: [3,2228,5,2,1,2219]
s: 2228
d: [2,1551,1,5,2,0,40,566,564]
s: 1551
d: [2,4617,6,13,4,2,44,4601,647]
d: [2,4635,6,10,3,3,21,4617,3774]
s: 4635
d: [2,1560,1,5,2,0,40,569,564]
s: 1560
d: [2,1768,1,5,2,2,44,1392,1575]
s: 1768
d: [2,4663,7,11,4,3,44,4635,584]
d: [2,4686,7,9,3,4,21,4663,1379]
d: [2,4687,7,7,2,5,21,4686,1931]
d: [2,4688,7,4,1,5,21,4687,3774]
a: [3,4692,7,4,1,4688]
s: 4692
r: 701
r: 3290
r: 3294
r: 3333
r: 3774
r: 3779
r: 4045
r: 4308
r: 4334
r: 4346
r: 4363
r: 4398
r: 4435
r: 4440
r: 4595
r: 4601
r: 4635
d: [2,4718,0,6,1,2,23,701,4692]
a: [3,4742,0,6,1,4718]
s: 4742
r: 3283
r: 4692
r: 4742
a: [3,3282,1,3,1,3279]
s: 3282
d: [2,4807,4,0,0,3,52,3602,3984,3282,606]
e: 4807
d: [2,4809,2,6,2,1,40,3282,606]
d: [2,4811,2,3,1,2,21,4809,3984]
d: [2,4812,2,0,0,3,21,4811,3602]
e: 4812
r: 2716
r: 3282
s: 4742
s: 701
s: 3283
s: 3290
s: 3294
s: 3333
d: [2,4745,1,7,2,1,44,4742,584]
d: [2,4760,1,4,1,2,21,4745,3597]
a: [3,4764,1,4,1,4760]
s: 4764
r: 701
r: 3283
r: 3286
r: 3336
r: 3597
r: 3602
r: 3984
r: 4320
r: 4340
r: 4352
r: 4383
r: 4414
r: 4419
r: 4742
s: 3774
d: [2,4746,1,11,3,1,44,4742,647]
d: [2,4766,1,8,2,2,21,4746,3597]
d: [2,4767,1,6,1,3,21,4766,1931]
a: [3,4771,1,6,1,4767]
s: 4771
s: 3779
d: [2,4836,1,8,2,2,44,4742,1575]
s: 4836
d: [2,4926,1,7,2,3,22,4836,4764]
s: 4926
s: 4308
s: 2716
s: 4334
d: [2,4851,1,17,4,1,44,660,4742]
d: [2,4872,1,15,3,2,21,4851,1931]
d: [2,4873,1,12,2,3,21,4872,3597]
s: 4873
d: [2,4956,1,11,2,4,22,4873,4764]
s: 4956
s: 4346
s: 4045
s: 4363
d: [2,4985,1,12,2,2,22,4045,4764]
s: 4985
d: [2,1902,1,5,2,0,40,549,622]
s: 1902
s: 4398
d: [2,1927,1,5,2,0,40,550,622]
s: 1927
s: 4435
d: [2,1951,1,5,2,0,40,559,622]
s: 1951
d: [2,1975,1,5,2,0,40,562,622]
s: 1975
s: 4440
d: [2,2130,1,5,2,0,40,617,592]
s: 2130
s: 4692
r: 3290
r: 3294
r: 3333
r: 3774
r: 3779
r: 4308
r: 4334
r: 4346
r: 4363
r: 4398
r: 4435
r: 4440
r: 4985
d: [2,2131,1,5,2,0,40,617,593]
s: 2131
d: [2,2587,1,13,3,0,44,586,585]
d: [2,2594,1,10,2,0,18,2587]
d: [2,2595,1,8,2,1,22,2594,971]
a: [3,2599,1,5,1,2595]
s: 2599
s: 4595
d: [2,1766,1,6,2,2,44,881,1575]
s: 1766
d: [2,4927,2,9,3,3,44,4926,1575]
s: 4927
d: [2,2199,1,6,2,0,40,551,622]
s: 2199
d: [2,2214,1,6,2,0,40,552,622]
s: 2214
d: [2,4940,2,12,6,3,44,656,4926]
d: [2,4944,2,10,5,4,21,4940,1956]
a: [3,4947,2,8,4,4944]
s: 4947
d: [2,5125,3,0,0,5,52,1956,2656,2716,1379,4947]
e: 5125
d: [2,5126,3,0,0,5,52,1615,2078,2083,746,4947]
e: 5126
d: [2,5127,3,6,3,2,40,4947,1379]
d: [2,5135,3,4,2,3,21,5127,2716]
d: [2,5136,3,2,1,4,21,5135,2656]
d: [2,5137,3,0,0,5,21,5136,1956]
e: 5137
d: [2,5129,3,6,3,2,40,4947,746]
d: [2,5139,3,4,2,3,21,5129,2083]
d: [2,5140,3,2,1,4,21,5139,2078]
d: [2,5141,3,0,0,5,21,5140,1615]
e: 5141
r: 4692
r: 4956
r: 4927
r: 4926
r: 4764
r: 4947
a: [3,4741,0,6,1,4718]
s: 4741
s: 3282
s: 701
s: 3286
a: [3,3291,1,3,1,3279]
s: 3291
s: 3294
s: 3333
s: 3336
a: [3,4763,1,4,1,4760]
s: 4763
s: 3597
a: [3,4770,1,6,1,4767]
s: 4770
s: 3602
d: [2,5147,1,8,2,2,44,4741,1575]
s: 5147
s: 4045
s: 3774
d: [2,2217,1,6,2,0,40,555,622]
s: 2217
s: 3779
d: [2,2232,1,6,2,0,40,556,622]
s: 2232
d: [2,2236,1,6,2,0,40,560,622]
s: 2236
a: [3,3983,3,3,1,3980]
s: 3983
d: [2,2239,1,6,2,0,40,563,622]
s: 2239
a: [3,4307,3,3,1,4304]
s: 4307
d: [2,2245,1,8,3,0,44,567,585]
d: [2,2274,1,6,2,0,21,2245,565]
s: 2274
d: [2,2365,1,6,2,0,40,616,622]
s: 2365
s: 4320
d: [2,2611,1,6,2,1,40,620,846]
s: 2611
s: 4334
a: [3,2660,1,6,2,2653]
s: 2660
a: [3,2712,1,6,2,2709]
s: 2712
d: [2,5289,2,6,3,1,44,2712,580]
d: [2,5300,2,4,2,2,21,5289,2716]
d: [2,5301,2,2,1,3,21,5300,1956]
d: [2,5302,2,0,0,4,21,5301,1931]
e: 5302
d: [2,5292,2,4,2,2,44,2712,906]
d: [2,5308,2,2,1,3,21,5292,2083]
d: [2,5309,2,0,0,4,21,5308,1615]
e: 5309
r: 2712
s: 4340
d: [2,3641,1,18,9,0,40,570,574]
d: [2,3664,1,14,7,0,18,3641]
a: [3,3671,1,6,3,3664]
s: 3671
d: [2,5313,2,0,0,4,52,2656,2716,1956,3671]
e: 5313
d: [2,5314,2,0,0,4,52,2078,2083,1615,3671]
e: 5314
d: [2,5319,2,4,2,2,40,3671,1956]
d: [2,5338,2,2,1,3,21,5319,2716]
d: [2,5339,2,0,0,4,21,5338,2656]
e: 5339
d: [2,5325,2,4,2,2,40,3671,1615]
d: [2,5341,2,2,1,3,21,5325,2083]
d: [2,5342,2,0,0,4,21,5341,2078]
e: 5342
r: 3671
s: 4346
d: [2,1764,1,7,3,1,44,1575,1575]
s: 1764
d: [2,2049,1,7,3,0,44,569,581]
s: 2049
s: 4352
d: [2,2069,1,10,4,0,44,598,581]
d: [2,2070,1,7,3,0,21,2069,579]
s: 2070
s: 4363
d: [2,2197,1,7,3,0,40,551,617]
s: 2197
d: [2,2212,1,7,3,0,40,552,617]
s: 2212
s: 4383
d: [2,2215,1,7,3,0,40,555,617]
s: 2215
s: 4398
d: [2,2230,1,7,3,0,40,556,617]
s: 2230
d: [2,2234,1,7,3,0,40,560,617]
s: 2234
s: 4414
d: [2,2237,1,7,3,0,40,563,617]
s: 2237
s: 4419
d: [2,2608,1,7,2,0,40,621,622]
s: 2608
d: [2,2756,1,11,4,0,44,646,580]
d: [2,2805,1,9,3,0,18,2756]
d: [2,2806,1,7,2,1,21,2805,1931]
s: 2806
s: 4435
d: [2,2760,1,11,4,0,44,646,583]
d: [2,2803,1,9,3,0,18,2760]
d: [2,2811,1,7,2,1,21,2803,1644]
s: 2811
s: 4440
d: [2,2764,1,11,4,0,44,646,580]
d: [2,2802,1,9,3,0,18,2764]
d: [2,2812,1,7,2,1,21,2802,1931]
s: 2812
d: [2,2768,1,11,4,0,44,646,583]
d: [2,2800,1,9,3,0,18,2768]
d: [2,2814,1,7,2,1,21,2800,1644]
s: 2814
a: [3,4691,7,4,1,4688]
s: 4691
d: [2,2774,1,11,4,0,44,580,646]
d: [2,2798,1,9,3,0,18,2774]
d: [2,2815,1,7,2,1,21,2798,1931]
s: 2815
a: [3,4594,4,6,1,4591]
s: 4594
d: [2,2778,1,11,4,0,44,583,646]
d: [2,2796,1,9,3,0,18,2778]
d: [2,2817,1,7,2,1,21,2796,1644]
s: 2817
d: [2,2784,1,11,4,0,44,580,646]
d: [2,2794,1,9,3,0,18,2784]
d: [2,2818,1,7,2,1,21,2794,1931]
s: 2818
d: [2,5187,2,8,2,2,44,4770,1575]
s: 5187
d: [2,2788,1,11,4,0,44,583,646]
d: [2,2792,1,9,3,0,18,2788]
d: [2,2820,1,7,2,1,21,2792,1644]
s: 2820
d: [2,5464,5,8,2,2,44,4594,1575]
s: 5464
d: [2,2822,1,11,4,0,44,647,584]
d: [2,2863,1,9,3,0,18,2822]
d: [2,2864,1,7,2,1,21,2863,1931]
s: 2864
d: [2,2830,1,11,4,0,44,647,584]
d: [2,2862,1,9,3,0,18,2830]
d: [2,2865,1,7,2,1,21,2862,1931]
s: 2865
d: [2,5191,2,13,4,2,44,5147,647]
d: [2,5200,2,10,3,3,21,5191,3597]
s: 5200
d: [2,2838,1,11,4,0,44,584,647]
d: [2,2861,1,9,3,0,18,2838]
d: [2,2866,1,7,2,1,21,2861,1931]
s: 2866
d: [2,5192,2,13,4,2,44,5147,647]
d: [2,5201,2,10,3,3,21,5192,3597]
s: 5201
d: [2,2848,1,11,4,0,44,584,647]
d: [2,2859,1,9,3,0,18,2848]
d: [2,2867,1,7,2,1,21,2859,1931]
s: 2867
d: [2,2874,1,13,6,0,44,629,585]
d: [2,2899,1,7,3,0,18,2874]
s: 2899
d: [2,5473,6,13,4,2,44,5464,647]
d: [2,5479,6,10,3,3,21,5473,3774]
s: 5479
d: [2,2900,1,13,6,0,44,630,585]
d: [2,2920,1,9,4,0,18,2900]
d: [2,2921,1,7,3,0,21,2920,565]
s: 2921
r: 3282
r: 3291
r: 3333
r: 3336
s: 3283
s: 3290
a: [3,3332,1,3,1,3329]
s: 3332
r: 3283
r: 3290
r: 3332
a: [3,2308,2,2,1,2287]
s: 2308
d: [2,5587,3,0,0,3,52,1439,2304,2308,608]
e: 5587
d: [2,5588,3,4,2,1,40,2308,608]
d: [2,5591,3,2,1,2,21,5588,2304]
d: [2,5592,3,0,0,3,21,5591,1439]
e: 5592
r: 2304
s: 3282
a: [3,2303,2,2,1,2287]
s: 2303
s: 3291
s: 3333
s: 3336
d: [2,3078,1,13,6,0,44,636,585]
d: [2,3098,1,9,4,0,18,3078]
d: [2,3099,1,7,3,0,21,3098,565]
s: 3099
d: [2,5602,3,0,0,3,52,2303,711,3282,3099]
e: 5602
d: [2,5603,2,4,2,1,40,3099,3282]
d: [2,5609,2,2,1,2,21,5603,711]
d: [2,5610,2,0,0,3,21,5609,2303]
e: 5610
r: 2303
r: 2308
r: 3282
r: 3291
r: 3333
r: 3336
s: 2304
s: 3283
s: 3290
a: [3,2307,2,2,1,2287]
s: 2307
s: 3332
a: [3,2321,2,2,1,2288]
s: 2321
d: [2,5619,3,0,0,3,52,1444,2317,2321,608]
e: 5619
d: [2,5620,3,4,2,1,40,2321,608]
d: [2,5623,3,2,1,2,21,5620,2317]
d: [2,5624,3,0,0,3,21,5623,1444]
e: 5624
r: 2317
r: 2078
r: 2083
r: 2321
r: 2716
r: 3290
r: 3332
r: 3983
r: 4307
r: 5479
r: 5464
r: 4594
r: 4691
r: 5147
r: 4741
r: 4763
r: 5187
r: 4770
s: 4742
s: 3290
s: 2317
s: 3333
a: [3,2320,2,2,1,2288]
s: 2320
s: 3336
s: 4764
r: 701
r: 3283
r: 3286
r: 3336
r: 3597
r: 3602
r: 4045
r: 4320
r: 4340
r: 4352
r: 4383
r: 4414
r: 4419
r: 4742
d: [2,5739,2,6,3,1,44,2899,4764]
d: [2,5744,2,4,2,2,21,5739,1379]
d: [2,5745,2,2,1,3,21,5744,1460]
d: [2,5746,2,0,0,4,21,5745,2299]
e: 5746
d: [2,3032,1,12,5,1,44,634,1016]
d: [2,3577,1,10,4,2,21,3032,1615]
d: [2,3608,1,15,5,1,44,658,1156]
d: [2,3631,1,12,4,1,18,3608]
d: [2,4422,2,6,2,1,40,2099,608]
d: [2,4448,2,3,1,1,40,2102,564]
d: [2,5168,1,19,4,1,44,660,701]
d: [2,5177,2,3,1,1,40,3286,565]
d: [2,5178,2,3,1,1,40,3286,564]
d: [2,5183,2,3,1,1,40,3294,565]
d: [2,5190,4,11,3,3,44,5147,4513]
d: [2,5195,4,8,2,3,18,5190]
d: [2,5196,4,3,1,3,19,5195]
d: [2,5528,2,4,2,1,40,2921,3282]
d: [2,5533,2,2,1,2,21,5528,711]
d: [2,5594,3,2,1,2,21,5588,1439]
d: [2,5614,2,4,2,1,40,3332,2921]
d: [2,5617,2,2,1,2,21,5614,706]
d: [2,5626,3,2,1,2,21,5620,1444]
d: [2,5613,2,4,2,1,40,3332,3099]
d: [2,5628,2,2,1,2,21,5613,706]
d: [2,4357,2,8,4,1,40,2026,572]
d: [2,5634,2,6,3,2,21,4357,2656]
d: [2,5635,2,4,2,3,21,5634,1615]
d: [2,5636,2,2,1,4,21,5635,1956]
d: [2,4356,2,8,4,1,40,2026,573]
d: [2,5638,2,6,3,2,21,4356,1615]
d: [2,5639,2,4,2,3,21,5638,1956]
d: [2,5167,1,14,5,1,44,656,701]
d: [2,5668,1,11,4,2,21,5167,3294]
d: [2,5669,1,8,3,3,21,5668,3286]
d: [2,5188,3,9,3,1,40,3602,595]
d: [2,5674,3,6,2,2,21,5188,3283]
d: [2,5685,1,7,2,1,44,584,4742]
d: [2,5705,1,4,1,2,21,5685,3597]
f: 5746,5705,5674,5669,5639,5636,5628,5626,5617,5610,5594,5533,5196,5183,5178,5177,5168,4492,4448,4445,4422,4403,3638,3631,3577,3395,2753,2351,2095,2022,1952,1919,1903,1741,1740,1610,1609,1608,1562,1502,1451,1431,1430,1406,1388,1381,1375,545,539,542,538,534,506,495,482,456,455,454,449,448,440,441,442,443
% Refutation found. Thanks to Tanya!
% SZS status Theorem for complex1__t1_complex1
% SZS output start Proof for complex1__t1_complex1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (k2_xcmplx_0(k3_square_1(X0),k3_square_1(X1)) = k6_numbers => k6_numbers = X0))) [input t1_complex1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (k2_xcmplx_0(k3_square_1(X0),k3_square_1(X1)) = k6_numbers => k6_numbers = X0))) [negated conjecture 1]
5. ! [X0] : (v1_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xcmplx_0(X1))) [input cc12_membered]
6. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xxreal_0(X1))) [input cc13_membered]
7. ! [X0] : (v3_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xreal_0(X1))) [input cc14_membered]
13. ! [X0] : (v6_membered(X0) => v5_membered(X0)) [input cc1_membered]
15. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xreal_0(X0)) [input cc1_xreal_0]
23. ! [X0] : (v5_membered(X0) => v4_membered(X0)) [input cc2_membered]
27. ! [X0] : (v4_membered(X0) => v3_membered(X0)) [input cc3_membered]
28. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
30. ! [X0] : (v3_membered(X0) => v2_membered(X0)) [input cc4_membered]
31. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
32. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) => (v2_xxreal_0(X0) & v1_xxreal_0(X0))) [input cc4_xxreal_0]
33. ! [X0] : (v3_membered(X0) => v1_membered(X0)) [input cc5_membered]
34. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
42. ! [X0] : (v1_xcmplx_0(X0) => k3_square_1(X0) = k3_xcmplx_0(X0,X0)) [input d1_square_1]
59. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & v2_xxreal_0(X0)) => v2_xxreal_0(k2_xcmplx_0(X0,X1))) [input fc11_xreal_0]
64. ! [X0] : ((v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => (~v2_xxreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0)))) [input fc16_xreal_0]
67. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & v2_xxreal_0(X0)) => v2_xxreal_0(k6_xcmplx_0(X0,X1))) [input fc19_xreal_0]
70. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
75. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v2_xxreal_0(k3_xcmplx_0(X1,X0))) [input fc24_xreal_0]
76. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v3_xxreal_0(k3_xcmplx_0(X0,X1))) [input fc25_xreal_0]
77. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k3_xcmplx_0(X0,X1))) [input fc26_xreal_0]
78. ! [X0] : (v1_xreal_0(X0) => v1_xreal_0(k3_square_1(X0))) [input fc2_square_1]
82. ! [X0] : (v1_xreal_0(X0) => (v1_xreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0)))) [input fc3_xreal_0]
87. v6_membered(k4_ordinal1) [input fc6_membered]
92. ! [X0,X1] : ((v1_xcmplx_0(X1) & ~v1_xboole_0(X1) & v1_xcmplx_0(X0) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k3_xcmplx_0(X0,X1))) [input fc8_xcmplx_0]
110. k5_numbers = k4_ordinal1 [input redefinition_k5_numbers]
111. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
175. k4_xcmplx_0(1) = k6_xcmplx_0(k4_xcmplx_0(2),k4_xcmplx_0(1)) [input rqRealDiff__k6_xcmplx_0__rm2_rm1_rm1]
177. 0 = k3_xcmplx_0(0,0) [input rqRealMult__k3_xcmplx_0__r0_r0_r0]
180. 0 = k3_xcmplx_0(0,k4_xcmplx_0(2)) [input rqRealMult__k3_xcmplx_0__r0_rm2_r0]
192. 1 = k4_xcmplx_0(k4_xcmplx_0(1)) [input rqRealNeg__k4_xcmplx_0__rm1_r1]
195. m1_subset_1(0,k1_numbers) & m1_subset_1(0,k5_numbers) & m2_subset_1(0,k1_numbers,k5_numbers) [input spc0_numerals]
198. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
201. m1_subset_1(2,k1_numbers) & m1_subset_1(2,k5_numbers) & m2_subset_1(2,k1_numbers,k5_numbers) & v2_xxreal_0(2) [input spc2_numerals]
203. ! [X0,X1,X2] : ((v1_xcmplx_0(X2) & v1_xcmplx_0(X1) & v1_xcmplx_0(X0)) => k2_xcmplx_0(k2_xcmplx_0(X0,X1),X2) = k2_xcmplx_0(X0,k2_xcmplx_0(X1,X2))) [input spc6_arithm]
207. ! [X0] : (v1_xcmplx_0(X0) => k2_xcmplx_0(X0,k6_numbers) = X0) [input t1_arithm]
208. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
209. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
212. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
222. ! [X0] : (v1_xreal_0(X0) => r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X0))) [input t63_xreal_1]
223. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
245. ? [X0] : (? [X1] : ((k6_numbers != X0 & k2_xcmplx_0(k3_square_1(X0),k3_square_1(X1)) = k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
246. ? [X0] : (? [X1] : (k6_numbers != X0 & k2_xcmplx_0(k3_square_1(X0),k3_square_1(X1)) = k6_numbers & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 245]
250. ! [X0] : (! [X1] : (v1_xxreal_0(X1) | ~m1_subset_1(X1,X0)) | ~v2_membered(X0)) [ennf transformation 6]
251. ! [X0] : (! [X1] : (v1_xcmplx_0(X1) | ~m1_subset_1(X1,X0)) | ~v1_membered(X0)) [ennf transformation 5]
254. ! [X0] : (v4_membered(X0) | ~v5_membered(X0)) [ennf transformation 23]
255. ! [X0] : (v3_membered(X0) | ~v4_membered(X0)) [ennf transformation 27]
258. ! [X0] : (v1_membered(X0) | ~v3_membered(X0)) [ennf transformation 33]
259. ! [X0] : (v2_membered(X0) | ~v3_membered(X0)) [ennf transformation 30]
260. ! [X0] : (! [X1] : (v1_xreal_0(X1) | ~m1_subset_1(X1,X0)) | ~v3_membered(X0)) [ennf transformation 7]
262. ! [X0] : (v5_membered(X0) | ~v6_membered(X0)) [ennf transformation 13]
265. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 31]
266. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 28]
267. ! [X0] : (v1_xreal_0(k3_square_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 78]
268. ! [X0] : (r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X0)) | ~v1_xreal_0(X0)) [ennf transformation 222]
269. ! [X0] : ((v1_xreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0))) | ~v1_xreal_0(X0)) [ennf transformation 82]
274. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 212]
275. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 274]
276. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 209]
277. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 276]
292. ! [X0] : (k2_xcmplx_0(X0,k6_numbers) = X0 | ~v1_xcmplx_0(X0)) [ennf transformation 207]
293. ! [X0] : (k3_square_1(X0) = k3_xcmplx_0(X0,X0) | ~v1_xcmplx_0(X0)) [ennf transformation 42]
296. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 223]
298. ! [X0] : (v1_xreal_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 15]
302. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | (v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0))) [ennf transformation 32]
303. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0)) [flattening 302]
308. ! [X0] : ((~v2_xxreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0))) | (~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 64]
309. ! [X0] : ((~v2_xxreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0))) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 308]
314. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 34]
315. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 314]
336. ! [X0,X1] : (~v1_xboole_0(k3_xcmplx_0(X0,X1)) | (~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0))) [ennf transformation 92]
337. ! [X0,X1] : (~v1_xboole_0(k3_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0)) [flattening 336]
338. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 77]
339. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 338]
346. ! [X0,X1] : (~v2_xxreal_0(k3_xcmplx_0(X1,X0)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 75]
347. ! [X0,X1] : (~v2_xxreal_0(k3_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 346]
350. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 76]
351. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 350]
390. ! [X0,X1] : (v2_xxreal_0(k2_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0))) [ennf transformation 59]
391. ! [X0,X1] : (v2_xxreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0)) [flattening 390]
394. ! [X0,X1] : (v2_xxreal_0(k6_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0))) [ennf transformation 67]
395. ! [X0,X1] : (v2_xxreal_0(k6_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0)) [flattening 394]
398. ! [X0,X1,X2] : (k2_xcmplx_0(k2_xcmplx_0(X0,X1),X2) = k2_xcmplx_0(X0,k2_xcmplx_0(X1,X2)) | (~v1_xcmplx_0(X2) | ~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0))) [ennf transformation 203]
399. ! [X0,X1,X2] : (k2_xcmplx_0(k2_xcmplx_0(X0,X1),X2) = k2_xcmplx_0(X0,k2_xcmplx_0(X1,X2)) | ~v1_xcmplx_0(X2) | ~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0)) [flattening 398]
402. ? [X0] : (? [X1] : (k6_numbers != X0 & k2_xcmplx_0(k3_square_1(X0),k3_square_1(X1)) = k6_numbers & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (k6_numbers != sK0 & k6_numbers = k2_xcmplx_0(k3_square_1(sK0),k3_square_1(X1)) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
403. ? [X1] : (k6_numbers != sK0 & k6_numbers = k2_xcmplx_0(k3_square_1(sK0),k3_square_1(X1)) & v1_xreal_0(X1)) => (k6_numbers != sK0 & k6_numbers = k2_xcmplx_0(k3_square_1(sK0),k3_square_1(sK1)) & v1_xreal_0(sK1)) [choice axiom]
404. (k6_numbers != sK0 & k6_numbers = k2_xcmplx_0(k3_square_1(sK0),k3_square_1(sK1)) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 246,403,402]
440. v1_xreal_0(sK0) [cnf transformation 404] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
441. v1_xreal_0(sK1) [cnf transformation 404] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
442. k6_numbers = k2_xcmplx_0(k3_square_1(sK0),k3_square_1(sK1)) [cnf transformation 404] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
443. k6_numbers != sK0 [cnf transformation 404] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
448. v6_membered(k4_ordinal1) [cnf transformation 87] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
449. v1_xboole_0(k1_xboole_0) [cnf transformation 70] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
454. k6_numbers = k1_xboole_0 [cnf transformation 111] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
455. k5_numbers = k4_ordinal1 [cnf transformation 110] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
456. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 208] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
482. 0 = k3_xcmplx_0(0,0) [cnf transformation 177] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
495. 1 = k4_xcmplx_0(k4_xcmplx_0(1)) [cnf transformation 192] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
506. 0 = k3_xcmplx_0(0,k4_xcmplx_0(2)) [cnf transformation 180] {a:0,w:6,wCS:2400,sine:2,thAx:0,allAx:1,thDist:-1}
534. k4_xcmplx_0(1) = k6_xcmplx_0(k4_xcmplx_0(2),k4_xcmplx_0(1)) [cnf transformation 175] {a:0,w:8,wCS:3200,sine:2,thAx:0,allAx:1,thDist:-1}
538. m1_subset_1(2,k1_numbers) [cnf transformation 201] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
539. v2_xxreal_0(1) [cnf transformation 198] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
542. m1_subset_1(1,k1_numbers) [cnf transformation 198] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
545. m1_subset_1(0,k1_numbers) [cnf transformation 195] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
549. ~m1_subset_1(X1,X0) | v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 250] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
550. ~m1_subset_1(X1,X0) | v1_xcmplx_0(X1) | ~v1_membered(X0) [cnf transformation 251] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
553. ~v5_membered(X0) | v4_membered(X0) [cnf transformation 254] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
554. ~v4_membered(X0) | v3_membered(X0) [cnf transformation 255] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
557. ~v3_membered(X0) | v1_membered(X0) [cnf transformation 258] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
558. ~v3_membered(X0) | v2_membered(X0) [cnf transformation 259] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
559. ~m1_subset_1(X1,X0) | v1_xreal_0(X1) | ~v3_membered(X0) [cnf transformation 260] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
561. ~v6_membered(X0) | v5_membered(X0) [cnf transformation 262] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
564. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 265] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
565. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 266] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
566. v1_xreal_0(k3_square_1(X0)) | ~v1_xreal_0(X0) [cnf transformation 267] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
567. r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X0)) | ~v1_xreal_0(X0) [cnf transformation 268] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
569. v1_xreal_0(k4_xcmplx_0(X0)) | ~v1_xreal_0(X0) [cnf transformation 269] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
572. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 275] {a:0,w:11,wCS:4400,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
573. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 277] {a:0,w:11,wCS:4400,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
584. k2_xcmplx_0(X0,k6_numbers) = X0 | ~v1_xcmplx_0(X0) [cnf transformation 292] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
585. k3_square_1(X0) = k3_xcmplx_0(X0,X0) | ~v1_xcmplx_0(X0) [cnf transformation 293] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
588. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 296] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
590. ~m1_subset_1(X0,k1_numbers) | v1_xreal_0(X0) [cnf transformation 298] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
595. ~v1_xxreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xboole_0(X0) [cnf transformation 303] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
601. ~v2_xxreal_0(k4_xcmplx_0(X0)) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 309] {a:0,w:7,wCS:2800,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
606. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 315] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
608. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 315] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
629. ~v1_xboole_0(k3_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) [cnf transformation 337] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
630. ~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 339] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
634. ~v2_xxreal_0(k3_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 347] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
636. ~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 351] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
656. v2_xxreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0) [cnf transformation 391] {a:0,w:12,wCS:4800,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
658. v2_xxreal_0(k6_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0) [cnf transformation 395] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
660. k2_xcmplx_0(k2_xcmplx_0(X0,X1),X2) = k2_xcmplx_0(X0,k2_xcmplx_0(X1,X2)) | ~v1_xcmplx_0(X2) | ~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0) [cnf transformation 399] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
694. 1 <=> k6_numbers = sK0 [avatar definition]
696. k6_numbers != sK0 <- (~1) [avatar component clause 694] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
697. ~1 [avatar split clause 443,694]
699. 2 <=> k6_numbers = k2_xcmplx_0(k3_square_1(sK0),k3_square_1(sK1)) [avatar definition]
701. k6_numbers = k2_xcmplx_0(k3_square_1(sK0),k3_square_1(sK1)) <- (2) [avatar component clause 699] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
702. 2 [avatar split clause 442,699]
704. 3 <=> v1_xreal_0(sK1) [avatar definition]
706. v1_xreal_0(sK1) <- (3) [avatar component clause 704] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
707. 3 [avatar split clause 441,704]
709. 4 <=> v1_xreal_0(sK0) [avatar definition]
711. v1_xreal_0(sK0) <- (4) [avatar component clause 709] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
712. 4 [avatar split clause 440,709]
734. 9 <=> v6_membered(k4_ordinal1) [avatar definition]
736. v6_membered(k4_ordinal1) <- (9) [avatar component clause 734] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
737. 9 [avatar split clause 448,734]
739. 10 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
741. v1_xboole_0(k1_xboole_0) <- (10) [avatar component clause 739] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
742. 10 [avatar split clause 449,739]
764. 15 <=> k6_numbers = k1_xboole_0 [avatar definition]
766. k6_numbers = k1_xboole_0 <- (15) [avatar component clause 764] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
767. 15 [avatar split clause 454,764]
769. 16 <=> k5_numbers = k4_ordinal1 [avatar definition]
771. k5_numbers = k4_ordinal1 <- (16) [avatar component clause 769] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
772. 16 [avatar split clause 455,769]
774. 17 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
776. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (17) [avatar component clause 774] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
777. 17 [avatar split clause 456,774]
904. 43 <=> 0 = k3_xcmplx_0(0,0) [avatar definition]
906. 0 = k3_xcmplx_0(0,0) <- (43) [avatar component clause 904] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
907. 43 [avatar split clause 482,904]
969. 56 <=> 1 = k4_xcmplx_0(k4_xcmplx_0(1)) [avatar definition]
971. 1 = k4_xcmplx_0(k4_xcmplx_0(1)) <- (56) [avatar component clause 969] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
972. 56 [avatar split clause 495,969]
1014. 65 <=> 0 = k3_xcmplx_0(0,k4_xcmplx_0(2)) [avatar definition]
1016. 0 = k3_xcmplx_0(0,k4_xcmplx_0(2)) <- (65) [avatar component clause 1014] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1017. 65 [avatar split clause 506,1014]
1154. 93 <=> k4_xcmplx_0(1) = k6_xcmplx_0(k4_xcmplx_0(2),k4_xcmplx_0(1)) [avatar definition]
1156. k4_xcmplx_0(1) = k6_xcmplx_0(k4_xcmplx_0(2),k4_xcmplx_0(1)) <- (93) [avatar component clause 1154] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1157. 93 [avatar split clause 534,1154]
1159. 94 <=> m1_subset_1(2,k1_numbers) [avatar definition]
1161. m1_subset_1(2,k1_numbers) <- (94) [avatar component clause 1159] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1162. 94 [avatar split clause 538,1159]
1179. 98 <=> m1_subset_1(1,k1_numbers) [avatar definition]
1181. m1_subset_1(1,k1_numbers) <- (98) [avatar component clause 1179] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1182. 98 [avatar split clause 542,1179]
1194. 101 <=> v2_xxreal_0(1) [avatar definition]
1196. v2_xxreal_0(1) <- (101) [avatar component clause 1194] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1197. 101 [avatar split clause 539,1194]
1199. 102 <=> m1_subset_1(0,k1_numbers) [avatar definition]
1201. m1_subset_1(0,k1_numbers) <- (102) [avatar component clause 1199] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1202. 102 [avatar split clause 545,1199]
1375. v1_xboole_0(k6_numbers) <- (10, 15) [backward demodulation 741,766] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
1377. 138 <=> v1_xboole_0(k6_numbers) [avatar definition]
1379. v1_xboole_0(k6_numbers) <- (138) [avatar component clause 1377] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1380. 138 | ~10 | ~15 [avatar split clause 1375,764,739,1377]
1381. v6_membered(k5_numbers) <- (9, 16) [backward demodulation 736,771] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
1383. 139 <=> v6_membered(k5_numbers) [avatar definition]
1385. v6_membered(k5_numbers) <- (139) [avatar component clause 1383] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1386. 139 | ~9 | ~16 [avatar split clause 1381,769,734,1383]
1387. m1_subset_1(k1_xboole_0,k5_numbers) <- (16, 17) [forward demodulation 776,771] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
1388. m1_subset_1(k6_numbers,k5_numbers) <- (15, 16, 17) [forward demodulation 1387,766] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
1390. 140 <=> m1_subset_1(k6_numbers,k5_numbers) [avatar definition]
1392. m1_subset_1(k6_numbers,k5_numbers) <- (140) [avatar component clause 1390] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1393. 140 | ~15 | ~16 | ~17 [avatar split clause 1388,774,769,764,1390]
1406. v5_membered(k5_numbers) <- (139) [resolution 561,1385] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1411. 143 <=> v5_membered(k5_numbers) [avatar definition]
1413. v5_membered(k5_numbers) <- (143) [avatar component clause 1411] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1414. 143 | ~139 [avatar split clause 1406,1383,1411]
1430. v1_xxreal_0(sK0) <- (4) [resolution 564,711] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1431. v1_xxreal_0(sK1) <- (3) [resolution 564,706] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1437. 147 <=> v1_xxreal_0(sK0) [avatar definition]
1439. v1_xxreal_0(sK0) <- (147) [avatar component clause 1437] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1440. 147 | ~4 [avatar split clause 1430,709,1437]
1442. 148 <=> v1_xxreal_0(sK1) [avatar definition]
1444. v1_xxreal_0(sK1) <- (148) [avatar component clause 1442] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1445. 148 | ~3 [avatar split clause 1431,704,1442]
1451. v1_xcmplx_0(sK0) <- (4) [resolution 565,711] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1458. 150 <=> v1_xcmplx_0(sK0) [avatar definition]
1460. v1_xcmplx_0(sK0) <- (150) [avatar component clause 1458] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1461. 150 | ~4 [avatar split clause 1451,709,1458]
1502. v4_membered(k5_numbers) <- (143) [resolution 1413,553] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1504. 158 <=> v4_membered(k5_numbers) [avatar definition]
1506. v4_membered(k5_numbers) <- (158) [avatar component clause 1504] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
1507. 158 | ~143 [avatar split clause 1502,1411,1504]
1562. v3_membered(k5_numbers) <- (158) [resolution 1506,554] {a:3,w:2,wCS:800,sine:255,thAx:0,allAx:4,thDist:-4}
1564. 167 <=> v3_membered(k5_numbers) [avatar definition]
1566. v3_membered(k5_numbers) <- (167) [avatar component clause 1564] {a:3,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:4,thDist:-4}
1567. 167 | ~158 [avatar split clause 1562,1504,1564]
1575. k6_numbers = X0 | ~v1_xboole_0(X0) <- (15) [forward demodulation 588,766] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1608. v1_xreal_0(0) <- (102) [resolution 590,1201] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
1609. v1_xreal_0(1) <- (98) [resolution 590,1181] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
1610. v1_xreal_0(2) <- (94) [resolution 590,1161] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
1613. 174 <=> v1_xreal_0(0) [avatar definition]
1615. v1_xreal_0(0) <- (174) [avatar component clause 1613] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1616. 174 | ~102 [avatar split clause 1608,1199,1613]
1618. 175 <=> v1_xreal_0(1) [avatar definition]
1620. v1_xreal_0(1) <- (175) [avatar component clause 1618] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1621. 175 | ~98 [avatar split clause 1609,1179,1618]
1623. 176 <=> v1_xreal_0(2) [avatar definition]
1625. v1_xreal_0(2) <- (176) [avatar component clause 1623] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1626. 176 | ~94 [avatar split clause 1610,1159,1623]
1740. v2_membered(k5_numbers) <- (167) [resolution 1566,558] {a:4,w:2,wCS:800,sine:255,thAx:0,allAx:5,thDist:-5}
1741. v1_membered(k5_numbers) <- (167) [resolution 1566,557] {a:4,w:2,wCS:800,sine:255,thAx:0,allAx:5,thDist:-5}
1743. 194 <=> v2_membered(k5_numbers) [avatar definition]
1745. v2_membered(k5_numbers) <- (194) [avatar component clause 1743] {a:4,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:5,thDist:-5}
1746. 194 | ~167 [avatar split clause 1740,1564,1743]
1748. 195 <=> v1_membered(k5_numbers) [avatar definition]
1751. 195 | ~167 [avatar split clause 1741,1564,1748]
1894. v1_xxreal_0(k6_numbers) | ~v2_membered(k5_numbers) <- (140) [resolution 549,1392] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1903. v1_xxreal_0(k6_numbers) <- (140, 194) [subsumption resolution 1894,1745] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1905. 217 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1907. v1_xxreal_0(k6_numbers) <- (217) [avatar component clause 1905] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1908. 217 | ~140 | ~194 [avatar split clause 1903,1743,1390,1905]
1919. v1_xcmplx_0(k6_numbers) | ~v1_membered(k5_numbers) <- (140) [resolution 550,1392] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1929. 220 <=> v1_xcmplx_0(k6_numbers) [avatar definition]
1931. v1_xcmplx_0(k6_numbers) <- (220) [avatar component clause 1929] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1932. ~195 | 220 | ~140 [avatar split clause 1919,1390,1929,1748]
1943. v1_xreal_0(k6_numbers) | ~v3_membered(k5_numbers) <- (140) [resolution 559,1392] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1952. v1_xreal_0(k6_numbers) <- (140, 167) [subsumption resolution 1943,1566] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1954. 223 <=> v1_xreal_0(k6_numbers) [avatar definition]
1956. v1_xreal_0(k6_numbers) <- (223) [avatar component clause 1954] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1957. 223 | ~140 | ~167 [avatar split clause 1952,1564,1390,1954]
2020. r1_xxreal_0(k6_numbers,0) | ~v1_xreal_0(0) <- (43) [superposition 567,906] {a:1,w:5,wCS:2000,sine:1,thAx:0,allAx:2,thDist:-2}
2022. r1_xxreal_0(k6_numbers,0) <- (43, 174) [subsumption resolution 2020,1615] {a:1,w:3,wCS:1200,sine:1,thAx:0,allAx:2,thDist:-2}
2024. 233 <=> r1_xxreal_0(k6_numbers,0) [avatar definition]
2026. r1_xxreal_0(k6_numbers,0) <- (233) [avatar component clause 2024] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2027. 233 | ~43 | ~174 [avatar split clause 2022,1613,904,2024]
2073. ~v2_xxreal_0(1) | ~v1_xreal_0(k4_xcmplx_0(1)) | v3_xxreal_0(k4_xcmplx_0(1)) <- (56) [superposition 601,971] {a:1,w:8,wCS:3200,sine:3,thAx:0,allAx:2,thDist:-2}
2077. 236 <=> v3_xxreal_0(0) [avatar definition]
2081. 237 <=> v2_xxreal_0(0) [avatar definition]
2091. 239 <=> v1_xreal_0(k4_xcmplx_0(2)) [avatar definition]
2093. ~v1_xreal_0(k4_xcmplx_0(2)) <- (~239) [avatar component clause 2091] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2095. ~v1_xreal_0(k4_xcmplx_0(1)) | v3_xxreal_0(k4_xcmplx_0(1)) <- (56, 101) [subsumption resolution 2073,1196] {a:1,w:6,wCS:2400,sine:3,thAx:0,allAx:2,thDist:-2}
2097. 240 <=> v3_xxreal_0(k4_xcmplx_0(1)) [avatar definition]
2099. v3_xxreal_0(k4_xcmplx_0(1)) <- (240) [avatar component clause 2097] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2101. 241 <=> v1_xreal_0(k4_xcmplx_0(1)) [avatar definition]
2102. v1_xreal_0(k4_xcmplx_0(1)) <- (241) [avatar component clause 2101] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2103. ~v1_xreal_0(k4_xcmplx_0(1)) <- (~241) [avatar component clause 2101] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2104. 240 | ~241 | ~56 | ~101 [avatar split clause 2095,1194,969,2101,2097]
2113. 242 <=> v2_xxreal_0(k4_xcmplx_0(2)) [avatar definition]
2122. 244 <=> v2_xxreal_0(k4_xcmplx_0(1)) [avatar definition]
2298. 263 <=> v1_xboole_0(sK0) [avatar definition]
2299. ~v1_xboole_0(sK0) <- (~263) [avatar component clause 2298] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2300. v1_xboole_0(sK0) <- (263) [avatar component clause 2298] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2302. 264 <=> v2_xxreal_0(sK0) [avatar definition]
2303. ~v2_xxreal_0(sK0) <- (~264) [avatar component clause 2302] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2306. 265 <=> v3_xxreal_0(sK0) [avatar definition]
2308. v3_xxreal_0(sK0) <- (265) [avatar component clause 2306] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2315. 267 <=> v2_xxreal_0(sK1) [avatar definition]
2319. 268 <=> v3_xxreal_0(sK1) [avatar definition]
2321. v3_xxreal_0(sK1) <- (268) [avatar component clause 2319] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2351. $false <- (~1, 15, 263) [unit resulting resolution 696,2300,1575] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
2355. 1 | ~15 | ~263 [avatar contradiction clause 2351]
2655. 290 <=> v3_xxreal_0(k6_numbers) [avatar definition]
2656. ~v3_xxreal_0(k6_numbers) <- (~290) [avatar component clause 2655] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2657. v3_xxreal_0(k6_numbers) <- (290) [avatar component clause 2655] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2714. 293 <=> v2_xxreal_0(k6_numbers) [avatar definition]
2748. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (290) [resolution 2657,606] {a:2,w:4,wCS:1600,sine:255,thAx:0,allAx:3,thDist:-3}
2752. ~v1_xxreal_0(k6_numbers) <- (138, 290) [subsumption resolution 2748,1379] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
2753. $false <- (138, 217, 290) [subsumption resolution 2752,1907] {a:2,w:0,sine:255,thAx:0,allAx:3,thDist:-3}
2754. ~138 | ~217 | ~290 [avatar contradiction clause 2753]
2874. ~v1_xboole_0(k3_square_1(X0)) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) | ~v1_xcmplx_0(X0) [superposition 629,585] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
2899. ~v1_xboole_0(k3_square_1(X0)) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) [duplicate literal removal 2874] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
2900. ~v3_xxreal_0(k3_square_1(X0)) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) | ~v1_xcmplx_0(X0) [superposition 630,585] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
2920. ~v3_xxreal_0(k3_square_1(X0)) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) | ~v1_xcmplx_0(X0) [duplicate literal removal 2900] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
2921. ~v3_xxreal_0(k3_square_1(X0)) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [subsumption resolution 2920,565] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
3032. ~v2_xxreal_0(0) | ~v1_xreal_0(0) | v3_xxreal_0(0) | ~v1_xreal_0(k4_xcmplx_0(2)) | v2_xxreal_0(k4_xcmplx_0(2)) <- (65) [superposition 634,1016] {a:1,w:12,wCS:4800,sine:2,thAx:0,allAx:2,thDist:-2}
3078. ~v3_xxreal_0(k3_square_1(X0)) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) | ~v1_xcmplx_0(X0) [superposition 636,585] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
3098. ~v3_xxreal_0(k3_square_1(X0)) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) | ~v1_xcmplx_0(X0) [duplicate literal removal 3078] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
3099. ~v3_xxreal_0(k3_square_1(X0)) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [subsumption resolution 3098,565] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
3281. 300 <=> v3_xxreal_0(k3_square_1(sK0)) [avatar definition]
3282. v3_xxreal_0(k3_square_1(sK0)) <- (300) [avatar component clause 3281] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
3283. ~v3_xxreal_0(k3_square_1(sK0)) <- (~300) [avatar component clause 3281] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
3285. 301 <=> v1_xreal_0(k3_square_1(sK0)) [avatar definition]
3286. v1_xreal_0(k3_square_1(sK0)) <- (301) [avatar component clause 3285] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
3287. ~v1_xreal_0(k3_square_1(sK0)) <- (~301) [avatar component clause 3285] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
3293. 303 <=> v1_xreal_0(k3_square_1(sK1)) [avatar definition]
3294. v1_xreal_0(k3_square_1(sK1)) <- (303) [avatar component clause 3293] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
3295. ~v1_xreal_0(k3_square_1(sK1)) <- (~303) [avatar component clause 3293] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
3331. 304 <=> v3_xxreal_0(k3_square_1(sK1)) [avatar definition]
3332. v3_xxreal_0(k3_square_1(sK1)) <- (304) [avatar component clause 3331] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
3335. 305 <=> v2_xxreal_0(k3_square_1(sK0)) [avatar definition]
3393. ~v1_xreal_0(sK0) <- (~301) [resolution 3287,566] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
3395. $false <- (4, ~301) [subsumption resolution 3393,711] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
3396. ~4 | 301 [avatar contradiction clause 3395]
3577. ~v2_xxreal_0(0) | v3_xxreal_0(0) | ~v1_xreal_0(k4_xcmplx_0(2)) | v2_xxreal_0(k4_xcmplx_0(2)) <- (65, 174) [subsumption resolution 3032,1615] {a:1,w:10,wCS:4000,sine:2,thAx:0,allAx:2,thDist:-2}
3578. 242 | ~239 | 236 | ~237 | ~65 | ~174 [avatar split clause 3577,1613,1014,2081,2077,2091,2113]
3595. 306 <=> v1_xcmplx_0(k3_square_1(sK0)) [avatar definition]
3597. v1_xcmplx_0(k3_square_1(sK0)) <- (306) [avatar component clause 3595] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
3600. 307 <=> v1_xxreal_0(k3_square_1(sK0)) [avatar definition]
3602. v1_xxreal_0(k3_square_1(sK0)) <- (307) [avatar component clause 3600] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
3608. v2_xxreal_0(k4_xcmplx_0(1)) | ~v1_xreal_0(k4_xcmplx_0(1)) | v2_xxreal_0(k4_xcmplx_0(1)) | ~v1_xreal_0(k4_xcmplx_0(2)) | ~v2_xxreal_0(k4_xcmplx_0(2)) <- (93) [superposition 658,1156] {a:1,w:15,wCS:6000,sine:2,thAx:0,allAx:2,thDist:-2}
3631. v2_xxreal_0(k4_xcmplx_0(1)) | ~v1_xreal_0(k4_xcmplx_0(1)) | ~v1_xreal_0(k4_xcmplx_0(2)) | ~v2_xxreal_0(k4_xcmplx_0(2)) <- (93) [duplicate literal removal 3608] {a:1,w:12,wCS:4800,sine:2,thAx:0,allAx:2,thDist:-2}
3634. ~242 | ~239 | ~241 | 244 | ~93 [avatar split clause 3631,1154,2122,2101,2091,2113]
3636. ~v1_xreal_0(sK1) <- (~303) [resolution 3295,566] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
3638. $false <- (3, ~303) [subsumption resolution 3636,706] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
3639. ~3 | 303 [avatar contradiction clause 3638]
3772. 316 <=> v1_xcmplx_0(k3_square_1(sK1)) [avatar definition]
3774. v1_xcmplx_0(k3_square_1(sK1)) <- (316) [avatar component clause 3772] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
3982. 328 <=> v1_xboole_0(k3_square_1(sK0)) [avatar definition]
4044. 329 <=> ! [X11] : (k2_xcmplx_0(k3_square_1(sK0),k2_xcmplx_0(k3_square_1(sK1),X11)) = k2_xcmplx_0(k6_numbers,X11) | ~v1_xcmplx_0(X11)) [avatar definition]
4045. k2_xcmplx_0(k3_square_1(sK0),k2_xcmplx_0(k3_square_1(sK1),X11)) = k2_xcmplx_0(k6_numbers,X11) | ~v1_xcmplx_0(X11) <- (329) [avatar component clause 4044] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
4356. ~v2_xxreal_0(k6_numbers) | v2_xxreal_0(0) | ~v1_xreal_0(0) | ~v1_xreal_0(k6_numbers) <- (233) [resolution 2026,573] {a:2,w:8,wCS:3200,sine:3,thAx:0,allAx:3,thDist:-3}
4357. ~v3_xxreal_0(0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(0) | ~v1_xreal_0(k6_numbers) <- (233) [resolution 2026,572] {a:2,w:8,wCS:3200,sine:3,thAx:0,allAx:3,thDist:-3}
4401. ~v1_xreal_0(2) <- (~239) [resolution 2093,569] {a:2,w:2,wCS:800,sine:2,thAx:0,allAx:3,thDist:-3}
4403. $false <- (176, ~239) [subsumption resolution 4401,1625] {a:2,w:0,sine:2,thAx:0,allAx:3,thDist:-3}
4404. ~176 | 239 [avatar contradiction clause 4403]
4422. ~v2_xxreal_0(k4_xcmplx_0(1)) | ~v1_xxreal_0(k4_xcmplx_0(1)) <- (240) [resolution 2099,608] {a:2,w:6,wCS:2400,sine:255,thAx:0,allAx:3,thDist:-3}
4425. 366 <=> v1_xxreal_0(k4_xcmplx_0(1)) [avatar definition]
4428. ~366 | ~244 | ~240 [avatar split clause 4422,2097,2122,4425]
4443. ~v1_xreal_0(1) <- (~241) [resolution 2103,569] {a:2,w:2,wCS:800,sine:2,thAx:0,allAx:3,thDist:-3}
4445. $false <- (175, ~241) [subsumption resolution 4443,1620] {a:2,w:0,sine:2,thAx:0,allAx:3,thDist:-3}
4446. ~175 | 241 [avatar contradiction clause 4445]
4448. v1_xxreal_0(k4_xcmplx_0(1)) <- (241) [resolution 2102,564] {a:2,w:3,wCS:1200,sine:255,thAx:0,allAx:3,thDist:-3}
4450. 366 | ~241 [avatar split clause 4448,2101,4425]
4475. k2_xcmplx_0(k3_square_1(sK0),k3_square_1(sK1)) = k2_xcmplx_0(k6_numbers,k6_numbers) | ~v1_xcmplx_0(k6_numbers) | ~v1_xcmplx_0(k3_square_1(sK1)) <- (329) [superposition 4045,584] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
4490. k2_xcmplx_0(k3_square_1(sK0),k3_square_1(sK1)) = k2_xcmplx_0(k6_numbers,k6_numbers) | ~v1_xcmplx_0(k3_square_1(sK1)) <- (220, 329) [subsumption resolution 4475,1931] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
4491. k2_xcmplx_0(k3_square_1(sK0),k3_square_1(sK1)) = k2_xcmplx_0(k6_numbers,k6_numbers) <- (220, 316, 329) [subsumption resolution 4490,3774] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
4492. k6_numbers = k2_xcmplx_0(k6_numbers,k6_numbers) <- (2, 220, 316, 329) [forward demodulation 4491,701] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
4494. 372 <=> k6_numbers = k2_xcmplx_0(k6_numbers,k6_numbers) [avatar definition]
4496. k6_numbers = k2_xcmplx_0(k6_numbers,k6_numbers) <- (372) [avatar component clause 4494] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
4497. 372 | ~2 | ~220 | ~316 | ~329 [avatar split clause 4492,4044,3772,1929,699,4494]
4513. k2_xcmplx_0(X0,X0) = X0 | ~v1_xboole_0(X0) <- (15, 372) [superposition 4496,1575] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
4740. 379 <=> k6_numbers = k2_xcmplx_0(k3_square_1(sK0),k6_numbers) [avatar definition]
4741. k6_numbers != k2_xcmplx_0(k3_square_1(sK0),k6_numbers) <- (~379) [avatar component clause 4740] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
4742. k6_numbers = k2_xcmplx_0(k3_square_1(sK0),k6_numbers) <- (379) [avatar component clause 4740] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
4762. 380 <=> k6_numbers = k3_square_1(sK0) [avatar definition]
4764. k6_numbers = k3_square_1(sK0) <- (380) [avatar component clause 4762] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
5147. k2_xcmplx_0(k3_square_1(sK0),X0) != X0 | ~v1_xboole_0(X0) <- (15, ~379) [superposition 4741,1575] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
5167. v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k3_square_1(sK1)) | v3_xxreal_0(k3_square_1(sK1)) | ~v1_xreal_0(k3_square_1(sK0)) | ~v2_xxreal_0(k3_square_1(sK0)) <- (2) [superposition 656,701] {a:1,w:14,wCS:1400,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
5168. k2_xcmplx_0(k6_numbers,X0) = k2_xcmplx_0(k3_square_1(sK0),k2_xcmplx_0(k3_square_1(sK1),X0)) | ~v1_xcmplx_0(X0) | ~v1_xcmplx_0(k3_square_1(sK1)) | ~v1_xcmplx_0(k3_square_1(sK0)) <- (2) [superposition 660,701] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
5176. ~306 | ~316 | 329 | ~2 [avatar split clause 5168,699,4044,3772,3595]
5177. v1_xcmplx_0(k3_square_1(sK0)) <- (301) [resolution 3286,565] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
5178. v1_xxreal_0(k3_square_1(sK0)) <- (301) [resolution 3286,564] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
5179. 306 | ~301 [avatar split clause 5177,3285,3595]
5180. 307 | ~301 [avatar split clause 5178,3285,3600]
5183. v1_xcmplx_0(k3_square_1(sK1)) <- (303) [resolution 3294,565] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
5185. 316 | ~303 [avatar split clause 5183,3293,3772]
5188. v3_xxreal_0(k3_square_1(sK0)) | v2_xxreal_0(k3_square_1(sK0)) | v1_xboole_0(k3_square_1(sK0)) <- (307) [resolution 3602,595] {a:3,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
5190. k3_square_1(sK0) != k3_square_1(sK0) | ~v1_xboole_0(k3_square_1(sK0)) | ~v1_xboole_0(k3_square_1(sK0)) <- (15, 372, ~379) [superposition 5147,4513] {a:4,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
5195. k3_square_1(sK0) != k3_square_1(sK0) | ~v1_xboole_0(k3_square_1(sK0)) <- (15, 372, ~379) [duplicate literal removal 5190] {a:4,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
5196. ~v1_xboole_0(k3_square_1(sK0)) <- (15, 372, ~379) [trivial inequality removal 5195] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
5199. ~328 | ~15 | ~372 | 379 [avatar split clause 5196,4740,4494,764,3982]
5528. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (300) [resolution 2921,3282] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5533. v3_xxreal_0(sK0) <- (4, 300) [subsumption resolution 5528,711] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5534. 265 | ~4 | ~300 [avatar split clause 5533,3281,709,2306]
5588. ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (265) [resolution 2308,608] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
5594. ~v2_xxreal_0(sK0) <- (147, 265) [subsumption resolution 5588,1439] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
5595. ~264 | ~147 | ~265 [avatar split clause 5594,2306,1437,2302]
5603. ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (300) [resolution 3099,3282] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5609. v2_xxreal_0(sK0) <- (4, 300) [subsumption resolution 5603,711] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5610. $false <- (4, ~264, 300) [subsumption resolution 5609,2303] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5611. ~4 | 264 | ~300 [avatar contradiction clause 5610]
5613. ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) <- (304) [resolution 3332,3099] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5614. ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) <- (304) [resolution 3332,2921] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5617. v3_xxreal_0(sK1) <- (3, 304) [subsumption resolution 5614,706] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5618. 268 | ~3 | ~304 [avatar split clause 5617,3331,704,2319]
5620. ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (268) [resolution 2321,608] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
5626. ~v2_xxreal_0(sK1) <- (148, 268) [subsumption resolution 5620,1444] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
5627. ~267 | ~148 | ~268 [avatar split clause 5626,2319,1442,2315]
5628. v2_xxreal_0(sK1) <- (3, 304) [subsumption resolution 5613,706] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5629. 267 | ~3 | ~304 [avatar split clause 5628,3331,704,2315]
5634. ~v3_xxreal_0(0) | ~v1_xreal_0(0) | ~v1_xreal_0(k6_numbers) <- (233, ~290) [subsumption resolution 4357,2656] {a:2,w:6,wCS:2400,sine:3,thAx:0,allAx:3,thDist:-3}
5635. ~v3_xxreal_0(0) | ~v1_xreal_0(k6_numbers) <- (174, 233, ~290) [subsumption resolution 5634,1615] {a:2,w:4,wCS:1600,sine:3,thAx:0,allAx:3,thDist:-3}
5636. ~v3_xxreal_0(0) <- (174, 223, 233, ~290) [subsumption resolution 5635,1956] {a:2,w:2,wCS:800,sine:3,thAx:0,allAx:3,thDist:-3}
5637. ~236 | ~174 | ~223 | ~233 | 290 [avatar split clause 5636,2655,2024,1954,1613,2077]
5638. ~v2_xxreal_0(k6_numbers) | v2_xxreal_0(0) | ~v1_xreal_0(k6_numbers) <- (174, 233) [subsumption resolution 4356,1615] {a:2,w:6,wCS:2400,sine:3,thAx:0,allAx:3,thDist:-3}
5639. ~v2_xxreal_0(k6_numbers) | v2_xxreal_0(0) <- (174, 223, 233) [subsumption resolution 5638,1956] {a:2,w:4,wCS:1600,sine:3,thAx:0,allAx:3,thDist:-3}
5640. 237 | ~293 | ~174 | ~223 | ~233 [avatar split clause 5639,2024,1954,1613,2714,2081]
5668. v2_xxreal_0(k6_numbers) | v3_xxreal_0(k3_square_1(sK1)) | ~v1_xreal_0(k3_square_1(sK0)) | ~v2_xxreal_0(k3_square_1(sK0)) <- (2, 303) [subsumption resolution 5167,3294] {a:1,w:11,wCS:1100,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
5669. v2_xxreal_0(k6_numbers) | v3_xxreal_0(k3_square_1(sK1)) | ~v2_xxreal_0(k3_square_1(sK0)) <- (2, 301, 303) [subsumption resolution 5668,3286] {a:1,w:8,wCS:800,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
5670. ~305 | 304 | 293 | ~2 | ~301 | ~303 [avatar split clause 5669,3293,3285,699,2714,3331,3335]
5674. v2_xxreal_0(k3_square_1(sK0)) | v1_xboole_0(k3_square_1(sK0)) <- (~300, 307) [subsumption resolution 5188,3283] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
5675. 328 | 305 | 300 | ~307 [avatar split clause 5674,3600,3281,3335,3982]
5685. k6_numbers = k3_square_1(sK0) | ~v1_xcmplx_0(k3_square_1(sK0)) <- (379) [superposition 584,4742] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
5705. k6_numbers = k3_square_1(sK0) <- (306, 379) [subsumption resolution 5685,3597] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
5706. 380 | ~306 | ~379 [avatar split clause 5705,4740,3595,4762]
5739. ~v1_xboole_0(k6_numbers) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (380) [superposition 2899,4764] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5744. ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (138, 380) [subsumption resolution 5739,1379] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5745. v1_xboole_0(sK0) <- (138, 150, 380) [subsumption resolution 5744,1460] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5746. $false <- (138, 150, ~263, 380) [subsumption resolution 5745,2299] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
5747. ~138 | ~150 | 263 | ~380 [avatar contradiction clause 5746]
5748. $false [avatar sat refutation 697,702,707,712,737,742,767,772,777,907,972,1017,1157,1162,1182,1197,1202,1380,1386,1393,1414,1440,1445,1461,1507,1567,1616,1621,1626,1746,1751,1908,1932,1957,2027,2104,2355,2754,3396,3578,3634,3639,4404,4428,4446,4450,4497,5176,5179,5180,5185,5199,5534,5595,5611,5618,5627,5629,5637,5640,5670,5675,5706,5747] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:4095,thDist:-4095}
% SZS output end Proof for complex1__t1_complex1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 227
% >>> Preprocessing
% Introduced skolems: 19
% Pure predicates: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 253
% Generated clauses: 3745
% Active clauses: 536
% Passive clauses: 985
% Final active clauses: 420
% Final passive clauses: 334
% Main loop iterations started: 566
% 
% >>> Simplifying Inferences
% Duplicate literals: 370
% Trivial inequalities: 4
% Fw subsumption resolutions: 684
% Fw demodulations: 20
% Bw demodulations: 60
% 
% >>> Deletion Inferences
% Simple tautologies: 331
% Equational tautologies: 2
% Forward subsumptions: 768
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Binary resolution: 544
% Unit resulting resolution: 16
% Factoring: 1
% Forward superposition: 818
% Backward superposition: 417
% Self superposition: 4
% 
% >>> AVATAR
% Split clauses: 255
% Split components: 675
% Sat splitting refutations: 32
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 654
% SAT solver unit clauses: 140
% SAT solver binary clauses: 164
% 
% Memory used [KB]: 8571
% Time elapsed: 0.036 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% SAT solver time: 0.001 s
% term sharing: 0.002 s
% congruence closure: 0.004 s
% backward demodulation index maintenance: 0.001 s
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.002 s
% backward subsumption: 0.002 s
% resolution: 0.002 s
% unit resulting resolution: 0.006 s ( own 0.005 s ) 
% superposition: 0.005 s ( own 0.004 s ) 
% passive container maintenance: 0.002 s
% other: 0.036 s ( own 0.009 s ) 

