i: [1,294,0,16,8,1,0,0] t1_gate_3
a: [3,302,0,2,1,294]
s: 302
a: [3,306,0,2,1,294]
s: 306
a: [3,310,0,2,1,294]
s: 310
a: [3,315,0,2,1,294]
s: 315
a: [3,318,0,2,1,294]
s: 318
a: [3,322,0,2,1,294]
s: 322
a: [3,326,0,2,1,294]
s: 326
a: [3,330,0,2,1,294]
s: 330
i: [1,38,0,4,2,1,0,0] t1_gate_3
a: [3,350,0,2,1,38]
s: 350
a: [3,355,0,2,1,38]
s: 355
i: [1,36,0,5,2,1,0,0] t1_gate_3
a: [3,365,0,2,1,36]
s: 365
a: [3,360,0,3,1,36]
s: 360
i: [1,34,0,6,2,1,0,0] t1_gate_3
a: [3,371,0,4,1,34]
s: 371
i: [1,26,0,6,2,1,0,0] t1_gate_3
a: [3,394,0,4,1,26]
s: 394
i: [1,32,0,7,2,1,0,0] t1_gate_3
a: [3,376,0,5,1,32]
s: 376
i: [1,30,0,7,2,1,0,0] t1_gate_3
a: [3,382,0,5,1,30]
s: 382
i: [1,24,0,7,2,1,0,0] t1_gate_3
a: [3,400,0,5,1,24]
s: 400
i: [1,22,0,7,2,1,0,0] t1_gate_3
a: [3,406,0,5,1,22]
s: 406
i: [1,28,0,8,2,1,0,0] t1_gate_3
a: [3,388,0,6,1,28]
s: 388
i: [1,20,0,8,2,1,0,0] t1_gate_3
a: [3,412,0,6,1,20]
s: 412
i: [1,295,0,5,2,0,0,1] fc2_gate_1
s: 295
r: 360
r: 365
a: [3,364,0,2,1,36]
s: 364
a: [3,361,0,3,1,36]
s: 361
i: [1,296,0,5,2,0,0,1] fc1_gate_1
s: 296
i: [1,298,0,6,2,0,0,1] fc4_gate_1
s: 298
d: [2,425,1,0,0,2,52,361,406,298]
e: 425
d: [2,426,1,0,0,2,52,361,412,298]
e: 426
d: [2,431,1,3,1,1,40,298,406]
d: [2,436,1,0,0,2,21,431,361]
e: 436
d: [2,433,1,3,1,1,40,298,412]
d: [2,438,1,0,0,2,21,433,361]
e: 438
r: 302
r: 310
r: 406
r: 412
a: [3,303,0,2,1,294]
s: 303
a: [3,311,0,2,1,294]
s: 311
a: [3,407,0,5,1,22]
s: 407
d: [2,416,1,2,1,1,40,295,360]
a: [3,419,1,2,1,416]
s: 419
a: [3,413,0,6,1,20]
s: 413
i: [1,297,0,7,1,0,0,255] commutativity_k2_gate_1
s: 297
r: 382
r: 388
r: 413
r: 407
r: 394
r: 371
r: 303
r: 311
r: 315
r: 318
r: 326
r: 361
r: 364
r: 400
r: 419
s: 303
s: 311
a: [3,314,0,2,1,294]
s: 314
s: 318
a: [3,327,0,2,1,294]
s: 327
s: 364
s: 361
s: 419
d: [2,452,0,4,1,1,23,394,297]
a: [3,479,0,4,1,452]
s: 479
d: [2,453,0,4,1,1,23,371,297]
a: [3,483,0,4,1,453]
s: 483
d: [2,503,1,0,0,2,52,364,483,298]
e: 503
d: [2,504,1,2,1,1,40,483,298]
d: [2,506,1,0,0,2,21,504,364]
e: 506
r: 318
r: 361
r: 364
r: 419
r: 479
a: [3,319,0,2,1,294]
s: 319
s: 365
a: [3,420,1,2,1,416]
s: 420
s: 360
a: [3,478,0,4,1,452]
s: 478
d: [2,432,1,3,1,1,40,298,388]
a: [3,443,1,3,1,432]
s: 443
d: [2,512,2,0,0,2,52,355,443,295]
e: 512
d: [2,513,2,2,1,1,40,443,295]
d: [2,515,2,0,0,2,21,513,355]
e: 515
r: 350
r: 355
r: 322
r: 376
a: [3,323,0,2,1,294]
s: 323
a: [3,351,0,2,1,38]
s: 351
a: [3,354,0,2,1,38]
s: 354
a: [3,377,0,5,1,32]
s: 377
a: [3,401,0,5,1,24]
s: 401
d: [2,448,0,5,1,1,23,382,297]
a: [3,459,0,5,1,448]
s: 459
d: [2,451,0,5,1,1,23,407,297]
a: [3,474,0,5,1,451]
s: 474
d: [2,449,0,6,1,1,23,388,297]
a: [3,464,0,6,1,449]
s: 464
d: [2,450,0,6,1,1,23,413,297]
a: [3,469,0,6,1,450]
s: 469
d: [2,520,1,3,1,1,40,464,298]
a: [3,524,1,3,1,520]
s: 524
d: [2,526,2,0,0,2,52,365,524,295]
e: 526
d: [2,527,2,2,1,1,40,524,295]
d: [2,529,2,0,0,2,21,527,365]
e: 529
r: 330
r: 464
r: 524
a: [3,331,0,2,1,294]
s: 331
a: [3,463,0,6,1,449]
s: 463
i: [1,299,0,8,3,0,0,1] fc3_gate_1
s: 299
r: 303
r: 469
s: 302
a: [3,523,1,3,1,520]
s: 523
a: [3,468,0,6,1,450]
s: 468
d: [2,534,1,6,2,1,40,299,469]
d: [2,542,1,3,1,2,21,534,360]
a: [3,545,1,3,1,542]
s: 545
d: [2,454,1,6,2,0,44,298,297]
s: 454
d: [2,556,2,0,0,2,52,354,483,454]
e: 556
d: [2,557,2,2,1,1,40,454,483]
d: [2,564,2,0,0,2,21,557,354]
e: 564
r: 311
r: 314
r: 474
r: 483
s: 310
s: 315
a: [3,484,0,4,1,453]
s: 484
a: [3,473,0,5,1,451]
s: 473
d: [2,571,1,0,0,2,52,351,473,298]
e: 571
d: [2,573,1,2,1,1,40,473,298]
d: [2,575,1,0,0,2,21,573,351]
e: 575
r: 306
r: 351
r: 354
r: 443
a: [3,383,0,5,1,30]
d: [2,582,0,5,1,1,22,383,297]
d: [2,583,0,0,0,2,21,582,459]
e: 583
r: 310
r: 459
r: 473
a: [3,307,0,2,1,294]
s: 307
s: 311
s: 351
s: 354
a: [3,442,1,3,1,432]
s: 442
d: [2,588,2,0,0,2,52,354,442,296]
e: 588
d: [2,589,2,2,1,1,40,442,296]
d: [2,591,2,0,0,2,21,589,354]
e: 591
r: 302
r: 323
r: 351
r: 354
r: 360
r: 365
r: 377
r: 420
r: 468
r: 545
r: 319
r: 327
r: 401
r: 478
r: 523
s: 303
s: 318
s: 322
s: 326
s: 350
s: 355
s: 364
s: 419
s: 361
s: 524
s: 479
s: 376
a: [3,546,1,3,1,542]
s: 546
s: 400
d: [2,609,2,0,0,2,52,546,400,454]
e: 609
d: [2,610,2,3,1,1,40,400,454]
d: [2,613,2,0,0,2,21,610,546]
e: 613
r: 322
r: 331
r: 350
r: 355
r: 376
r: 442
r: 463
r: 546
r: 326
r: 400
s: 323
s: 327
s: 331
s: 350
s: 355
s: 442
s: 377
d: [2,622,2,0,0,3,52,355,524,377,299]
e: 622
d: [2,623,1,5,2,1,40,377,299]
d: [2,625,1,2,1,2,21,623,524]
d: [2,626,1,0,0,3,21,625,355]
e: 626
r: 307
r: 318
r: 350
r: 355
r: 361
r: 364
r: 419
r: 442
r: 479
r: 524
s: 306
s: 319
s: 351
s: 354
s: 365
s: 360
s: 420
s: 478
s: 443
s: 401
s: 459
s: 523
s: 474
s: 546
d: [2,642,2,0,0,2,52,351,546,296]
e: 642
d: [2,643,2,2,1,1,40,546,296]
d: [2,645,2,0,0,2,21,643,351]
e: 645
r: 303
r: 306
r: 459
r: 546
s: 302
s: 307
s: 545
a: [3,458,0,5,1,448]
s: 458
d: [2,652,2,0,0,3,52,365,443,458,299]
e: 652
d: [2,653,1,5,2,1,40,458,299]
d: [2,655,1,2,1,2,21,653,443]
d: [2,656,1,0,0,3,21,655,365]
e: 656
r: 302
r: 307
r: 327
r: 331
r: 360
r: 365
r: 401
r: 420
r: 458
r: 523
s: 303
s: 306
s: 326
s: 330
s: 364
s: 419
s: 361
s: 400
s: 524
s: 459
d: [2,668,1,0,0,2,52,364,459,298]
e: 668
d: [2,670,1,2,1,1,40,459,298]
d: [2,672,1,0,0,2,21,670,364]
e: 672
r: 306
r: 315
r: 323
r: 326
r: 330
r: 351
r: 354
r: 361
r: 364
r: 377
r: 400
r: 419
r: 443
r: 459
r: 484
r: 524
r: 545
r: 311
r: 474
s: 307
s: 310
s: 314
s: 322
s: 327
s: 331
s: 350
s: 355
s: 365
s: 360
s: 420
s: 483
s: 376
i: [1,231,0,16,8,1,0,0] t1_gate_3
i: [1,87,0,16,8,1,0,0] t1_gate_3
i: [1,51,0,16,8,1,0,0] t1_gate_3
i: [1,42,0,16,8,1,0,0] t1_gate_3
i: [1,37,0,4,2,1,0,0] t1_gate_3
i: [1,35,0,5,2,1,0,0] t1_gate_3
i: [1,33,0,6,2,1,0,0] t1_gate_3
i: [1,31,0,7,2,1,0,0] t1_gate_3
i: [1,29,0,7,2,1,0,0] t1_gate_3
i: [1,27,0,8,2,1,0,0] t1_gate_3
i: [1,25,0,6,2,1,0,0] t1_gate_3
i: [1,23,0,7,2,1,0,0] t1_gate_3
i: [1,21,0,7,2,1,0,0] t1_gate_3
i: [1,19,0,8,2,1,0,0] t1_gate_3
d: [2,430,1,3,1,1,40,298,382]
d: [2,428,1,2,1,1,40,298,400]
d: [2,427,1,2,1,1,40,298,394]
d: [2,500,1,2,1,1,40,361,296]
d: [2,429,1,2,1,1,40,298,376]
d: [2,535,1,6,2,1,40,299,463]
d: [2,548,1,3,1,2,21,535,443]
d: [2,553,1,3,1,1,40,468,298]
d: [2,568,0,4,1,1,22,371,297]
d: [2,555,2,2,1,1,40,545,295]
d: [2,586,0,5,1,1,22,407,297]
d: [2,570,1,4,2,1,40,484,299]
d: [2,552,2,2,1,1,40,523,296]
d: [2,538,1,5,2,1,40,299,401]
d: [2,600,0,4,1,1,22,394,297]
d: [2,605,1,2,1,1,40,479,298]
a: [3,395,0,4,1,26]
d: [2,629,0,4,1,1,22,395,297]
d: [2,647,0,6,1,1,22,412,297]
d: [2,658,0,6,1,1,22,413,297]
d: [2,660,0,6,1,1,22,388,297]
d: [2,662,0,5,1,1,22,382,297]
d: [2,635,1,4,2,1,40,478,299]
d: [2,641,1,5,2,1,40,474,299]
a: [3,389,0,6,1,28]
d: [2,676,0,6,1,1,22,389,297]
d: [2,678,0,5,1,1,22,383,297]
a: [3,370,0,4,1,34]
d: [2,680,0,4,1,1,22,370,297]
d: [2,682,0,5,1,1,22,406,297]
d: [2,684,1,2,1,1,40,360,295]
d: [2,688,2,3,1,1,40,376,454]
f: 688,684,682,680,678,676,641,635,672,662,660,658,656,647,645,629,623,610,605,600,538,552,570,589,586,555,573,568,564,553,548,542,529,520,429,513,504,500,427,428,430,432,438,436,19,20,21,22,23,24,25,26,27,28,29,30,31,32,33,34,35,36,37,38,42,51,87,231
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_3__t1_gate_3
% SZS output start Proof for gate_3__t1_gate_3
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ~(~(~(v1_xboole_0(X4) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X5))) & ~(v1_xboole_0(X11) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k1_gate_1(X9))) & ~(v1_xboole_0(k1_gate_1(X9)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X11,X10))) & ~(v1_xboole_0(k2_gate_1(X11,X10)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k2_gate_1(X9,X8))) & ~(v1_xboole_0(k2_gate_1(X9,X8)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) & ~v1_xboole_0(X0))) [input t1_gate_3]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ~(~(~(v1_xboole_0(X4) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X5))) & ~(v1_xboole_0(X11) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k1_gate_1(X9))) & ~(v1_xboole_0(k1_gate_1(X9)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X11,X10))) & ~(v1_xboole_0(k2_gate_1(X11,X10)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k2_gate_1(X9,X8))) & ~(v1_xboole_0(k2_gate_1(X9,X8)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) & ~v1_xboole_0(X0))) [negated conjecture 1]
3. ! [X0,X1] : k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [input commutativity_k2_gate_1]
6. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
7. ! [X0] : (~v1_xboole_0(X0) => v1_xboole_0(k1_gate_1(X0))) [input fc2_gate_1]
8. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_gate_1(X0,X1))) [input fc3_gate_1]
9. ! [X0,X1] : (v1_xboole_0(X0) => v1_xboole_0(k2_gate_1(X0,X1))) [input fc4_gate_1]
10. ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ~(~(~(v1_xboole_0(X4) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X5))) & ~(v1_xboole_0(X11) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k1_gate_1(X9))) & ~(v1_xboole_0(k1_gate_1(X9)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X11,X10))) & ~(v1_xboole_0(k2_gate_1(X11,X10)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k2_gate_1(X9,X8))) & ~(v1_xboole_0(k2_gate_1(X9,X8)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) & ~v1_xboole_0(X0))) [flattening 2]
11. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (((v1_xboole_0(X4) & ~v1_xboole_0(X2)) | (v1_xboole_0(X2) & ~v1_xboole_0(X4)) | (v1_xboole_0(X6) & ~v1_xboole_0(X3)) | (v1_xboole_0(X3) & ~v1_xboole_0(X6)) | (v1_xboole_0(X7) & ~v1_xboole_0(X1)) | (v1_xboole_0(X1) & ~v1_xboole_0(X7)) | (v1_xboole_0(X5) & ~v1_xboole_0(X0)) | (v1_xboole_0(X0) & ~v1_xboole_0(X5))) & (~v1_xboole_0(X11) | v1_xboole_0(X8)) & (~v1_xboole_0(X8) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k1_gate_1(X9))) & (~v1_xboole_0(k1_gate_1(X9)) | v1_xboole_0(X10)) & (~v1_xboole_0(X7) | v1_xboole_0(k2_gate_1(X11,X10))) & (~v1_xboole_0(k2_gate_1(X11,X10)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & (~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k2_gate_1(X9,X8))) & (~v1_xboole_0(k2_gate_1(X9,X8)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & (~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) | v1_xboole_0(X0))) [ennf transformation 10]
12. ! [X0] : (v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0)) [ennf transformation 7]
13. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 6]
14. ! [X0,X1] : (v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 9]
15. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 8]
16. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 15]
17. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (((v1_xboole_0(X4) & ~v1_xboole_0(X2)) | (v1_xboole_0(X2) & ~v1_xboole_0(X4)) | (v1_xboole_0(X6) & ~v1_xboole_0(X3)) | (v1_xboole_0(X3) & ~v1_xboole_0(X6)) | (v1_xboole_0(X7) & ~v1_xboole_0(X1)) | (v1_xboole_0(X1) & ~v1_xboole_0(X7)) | (v1_xboole_0(X5) & ~v1_xboole_0(X0)) | (v1_xboole_0(X0) & ~v1_xboole_0(X5))) & (~v1_xboole_0(X11) | v1_xboole_0(X8)) & (~v1_xboole_0(X8) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k1_gate_1(X9))) & (~v1_xboole_0(k1_gate_1(X9)) | v1_xboole_0(X10)) & (~v1_xboole_0(X7) | v1_xboole_0(k2_gate_1(X11,X10))) & (~v1_xboole_0(k2_gate_1(X11,X10)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & (~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k2_gate_1(X9,X8))) & (~v1_xboole_0(k2_gate_1(X9,X8)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & (~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) | v1_xboole_0(X0))) => (((v1_xboole_0(sK4) & ~v1_xboole_0(sK2)) | (v1_xboole_0(sK2) & ~v1_xboole_0(sK4)) | (v1_xboole_0(sK6) & ~v1_xboole_0(sK3)) | (v1_xboole_0(sK3) & ~v1_xboole_0(sK6)) | (v1_xboole_0(sK7) & ~v1_xboole_0(sK1)) | (v1_xboole_0(sK1) & ~v1_xboole_0(sK7)) | (v1_xboole_0(sK5) & ~v1_xboole_0(sK0)) | (v1_xboole_0(sK0) & ~v1_xboole_0(sK5))) & (~v1_xboole_0(sK11) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK8) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k1_gate_1(sK9))) & (~v1_xboole_0(k1_gate_1(sK9)) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK7) | v1_xboole_0(k2_gate_1(sK11,sK10))) & (~v1_xboole_0(k2_gate_1(sK11,sK10)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10)))) & (~v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k2_gate_1(sK9,sK8))) & (~v1_xboole_0(k2_gate_1(sK9,sK8)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8)))) & (~v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) | v1_xboole_0(sK0))) [choice axiom]
18. ((v1_xboole_0(sK4) & ~v1_xboole_0(sK2)) | (v1_xboole_0(sK2) & ~v1_xboole_0(sK4)) | (v1_xboole_0(sK6) & ~v1_xboole_0(sK3)) | (v1_xboole_0(sK3) & ~v1_xboole_0(sK6)) | (v1_xboole_0(sK7) & ~v1_xboole_0(sK1)) | (v1_xboole_0(sK1) & ~v1_xboole_0(sK7)) | (v1_xboole_0(sK5) & ~v1_xboole_0(sK0)) | (v1_xboole_0(sK0) & ~v1_xboole_0(sK5))) & (~v1_xboole_0(sK11) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK8) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k1_gate_1(sK9))) & (~v1_xboole_0(k1_gate_1(sK9)) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK7) | v1_xboole_0(k2_gate_1(sK11,sK10))) & (~v1_xboole_0(k2_gate_1(sK11,sK10)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10)))) & (~v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k2_gate_1(sK9,sK8))) & (~v1_xboole_0(k2_gate_1(sK9,sK8)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8)))) & (~v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) | v1_xboole_0(sK0)) [skolemisation 11,17]
19. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) | v1_xboole_0(sK0) [cnf transformation 18] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
20. ~v1_xboole_0(sK0) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) [cnf transformation 18] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
21. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) | v1_xboole_0(sK1) [cnf transformation 18] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
22. ~v1_xboole_0(sK1) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) [cnf transformation 18] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
23. ~v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) | v1_xboole_0(sK2) [cnf transformation 18] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
24. ~v1_xboole_0(sK2) | v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) [cnf transformation 18] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
25. ~v1_xboole_0(k2_gate_1(sK9,sK8)) | v1_xboole_0(sK3) [cnf transformation 18] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
26. ~v1_xboole_0(sK3) | v1_xboole_0(k2_gate_1(sK9,sK8)) [cnf transformation 18] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) | v1_xboole_0(sK4) [cnf transformation 18] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. ~v1_xboole_0(sK4) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) [cnf transformation 18] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) | v1_xboole_0(sK5) [cnf transformation 18] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. ~v1_xboole_0(sK5) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) [cnf transformation 18] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) | v1_xboole_0(sK6) [cnf transformation 18] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. ~v1_xboole_0(sK6) | v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) [cnf transformation 18] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. ~v1_xboole_0(k2_gate_1(sK11,sK10)) | v1_xboole_0(sK7) [cnf transformation 18] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. ~v1_xboole_0(sK7) | v1_xboole_0(k2_gate_1(sK11,sK10)) [cnf transformation 18] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. ~v1_xboole_0(k1_gate_1(sK9)) | v1_xboole_0(sK10) [cnf transformation 18] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. ~v1_xboole_0(sK10) | v1_xboole_0(k1_gate_1(sK9)) [cnf transformation 18] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. ~v1_xboole_0(sK8) | v1_xboole_0(sK11) [cnf transformation 18] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. ~v1_xboole_0(sK11) | v1_xboole_0(sK8) [cnf transformation 18] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK7) | v1_xboole_0(sK5) | v1_xboole_0(sK0) [cnf transformation 18] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK6) | v1_xboole_0(sK7) | v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK5) [cnf transformation 18] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
87. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK4) | v1_xboole_0(sK6) | v1_xboole_0(sK3) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK7) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK5) [cnf transformation 18] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
231. v1_xboole_0(sK4) | v1_xboole_0(sK2) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK7) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK5) [cnf transformation 18] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
295. v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0) [cnf transformation 12] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
296. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 13] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
297. k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
298. v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 14] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
299. ~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 16] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
301. 1 <=> v1_xboole_0(sK0) [avatar definition]
305. 2 <=> v1_xboole_0(sK5) [avatar definition]
309. 3 <=> v1_xboole_0(sK1) [avatar definition]
313. 4 <=> v1_xboole_0(sK7) [avatar definition]
317. 5 <=> v1_xboole_0(sK3) [avatar definition]
321. 6 <=> v1_xboole_0(sK6) [avatar definition]
325. 7 <=> v1_xboole_0(sK2) [avatar definition]
329. 8 <=> v1_xboole_0(sK4) [avatar definition]
339. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | 7 | 8 [avatar split clause 231,329,325,317,321,309,313,301,305]
343. ~2 | ~1 | ~4 | ~3 | 5 | 6 | ~8 | ~7 [avatar split clause 87,325,329,321,317,309,313,301,305]
345. ~2 | ~1 | 3 | 4 | ~6 | ~5 | ~8 | ~7 [avatar split clause 51,325,329,317,321,313,309,301,305]
346. 1 | 2 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 [avatar split clause 42,325,329,317,321,309,313,305,301]
349. 9 <=> v1_xboole_0(sK8) [avatar definition]
351. v1_xboole_0(sK8) <- (9) [avatar component clause 349] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
353. 10 <=> v1_xboole_0(sK11) [avatar definition]
354. v1_xboole_0(sK11) <- (10) [avatar component clause 353] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
356. 9 | ~10 [avatar split clause 38,353,349]
357. 10 | ~9 [avatar split clause 37,349,353]
359. 11 <=> v1_xboole_0(k1_gate_1(sK9)) [avatar definition]
360. ~v1_xboole_0(k1_gate_1(sK9)) <- (~11) [avatar component clause 359] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
361. v1_xboole_0(k1_gate_1(sK9)) <- (11) [avatar component clause 359] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
363. 12 <=> v1_xboole_0(sK10) [avatar definition]
364. v1_xboole_0(sK10) <- (12) [avatar component clause 363] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
365. ~v1_xboole_0(sK10) <- (~12) [avatar component clause 363] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
366. 11 | ~12 [avatar split clause 36,363,359]
367. 12 | ~11 [avatar split clause 35,359,363]
369. 13 <=> v1_xboole_0(k2_gate_1(sK11,sK10)) [avatar definition]
370. ~v1_xboole_0(k2_gate_1(sK11,sK10)) <- (~13) [avatar component clause 369] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
371. v1_xboole_0(k2_gate_1(sK11,sK10)) <- (13) [avatar component clause 369] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
372. 13 | ~4 [avatar split clause 34,313,369]
373. 4 | ~13 [avatar split clause 33,369,313]
375. 14 <=> v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) [avatar definition]
376. ~v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) <- (~14) [avatar component clause 375] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
377. v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) <- (14) [avatar component clause 375] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
378. 14 | ~6 [avatar split clause 32,321,375]
379. 6 | ~14 [avatar split clause 31,375,321]
381. 15 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) [avatar definition]
382. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) <- (~15) [avatar component clause 381] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
383. v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) <- (15) [avatar component clause 381] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
384. 15 | ~2 [avatar split clause 30,305,381]
385. 2 | ~15 [avatar split clause 29,381,305]
387. 16 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) [avatar definition]
388. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) <- (~16) [avatar component clause 387] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
389. v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) <- (16) [avatar component clause 387] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
390. 16 | ~8 [avatar split clause 28,329,387]
391. 8 | ~16 [avatar split clause 27,387,329]
393. 17 <=> v1_xboole_0(k2_gate_1(sK9,sK8)) [avatar definition]
394. ~v1_xboole_0(k2_gate_1(sK9,sK8)) <- (~17) [avatar component clause 393] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
395. v1_xboole_0(k2_gate_1(sK9,sK8)) <- (17) [avatar component clause 393] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
396. 17 | ~5 [avatar split clause 26,317,393]
397. 5 | ~17 [avatar split clause 25,393,317]
399. 18 <=> v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) [avatar definition]
400. ~v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) <- (~18) [avatar component clause 399] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
401. v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) <- (18) [avatar component clause 399] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
402. 18 | ~7 [avatar split clause 24,325,399]
403. 7 | ~18 [avatar split clause 23,399,325]
405. 19 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) [avatar definition]
406. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) <- (~19) [avatar component clause 405] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
407. v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) <- (19) [avatar component clause 405] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
408. 19 | ~3 [avatar split clause 22,309,405]
409. 3 | ~19 [avatar split clause 21,405,309]
411. 20 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) [avatar definition]
412. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) <- (~20) [avatar component clause 411] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
413. v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) <- (20) [avatar component clause 411] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
414. 20 | ~1 [avatar split clause 20,301,411]
415. 1 | ~20 [avatar split clause 19,411,301]
418. 21 <=> v1_xboole_0(sK9) [avatar definition]
427. ~v1_xboole_0(sK9) <- (~17) [resolution 298,394] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
428. ~v1_xboole_0(sK9) <- (~18) [resolution 298,400] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
429. ~v1_xboole_0(sK11) <- (~14) [resolution 298,376] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
430. ~v1_xboole_0(k1_gate_1(sK11)) <- (~15) [resolution 298,382] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
431. ~v1_xboole_0(k1_gate_1(sK9)) <- (~19) [resolution 298,406] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
432. ~v1_xboole_0(k1_gate_1(sK11)) <- (~16) [resolution 298,388] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
433. ~v1_xboole_0(k1_gate_1(sK9)) <- (~20) [resolution 298,412] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
436. $false <- (11, ~19) [subsumption resolution 431,361] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
437. ~11 | 19 [avatar contradiction clause 436]
438. $false <- (11, ~20) [subsumption resolution 433,361] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
439. ~11 | 20 [avatar contradiction clause 438]
441. 22 <=> v1_xboole_0(k1_gate_1(sK11)) [avatar definition]
442. v1_xboole_0(k1_gate_1(sK11)) <- (22) [avatar component clause 441] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
443. ~v1_xboole_0(k1_gate_1(sK11)) <- (~22) [avatar component clause 441] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
444. ~22 | 16 [avatar split clause 432,387,441]
445. ~22 | 15 [avatar split clause 430,381,441]
446. ~21 | 18 [avatar split clause 428,399,418]
447. ~21 | 17 [avatar split clause 427,393,418]
454. v1_xboole_0(k2_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 298,297] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
457. 23 <=> v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) [avatar definition]
458. v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) <- (23) [avatar component clause 457] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
459. ~v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) <- (~23) [avatar component clause 457] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
462. 24 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) [avatar definition]
463. v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) <- (24) [avatar component clause 462] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
464. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) <- (~24) [avatar component clause 462] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
467. 25 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) [avatar definition]
468. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) <- (~25) [avatar component clause 467] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
469. v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) <- (25) [avatar component clause 467] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
472. 26 <=> v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) [avatar definition]
473. ~v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) <- (~26) [avatar component clause 472] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
474. v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) <- (26) [avatar component clause 472] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
477. 27 <=> v1_xboole_0(k2_gate_1(sK8,sK9)) [avatar definition]
478. v1_xboole_0(k2_gate_1(sK8,sK9)) <- (27) [avatar component clause 477] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
479. ~v1_xboole_0(k2_gate_1(sK8,sK9)) <- (~27) [avatar component clause 477] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
482. 28 <=> v1_xboole_0(k2_gate_1(sK10,sK11)) [avatar definition]
483. ~v1_xboole_0(k2_gate_1(sK10,sK11)) <- (~28) [avatar component clause 482] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
484. v1_xboole_0(k2_gate_1(sK10,sK11)) <- (28) [avatar component clause 482] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
500. ~v1_xboole_0(sK9) <- (11) [resolution 361,296] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
501. ~21 | ~11 [avatar split clause 500,359,418]
504. ~v1_xboole_0(sK10) <- (~28) [resolution 483,298] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
508. ~12 | 28 [avatar split clause 504,482,363]
513. v1_xboole_0(sK11) <- (~22) [resolution 443,295] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
517. 10 | 22 [avatar split clause 513,441,353]
518. ~10 | 14 [avatar split clause 429,375,353]
520. ~v1_xboole_0(k1_gate_1(sK10)) <- (~24) [resolution 464,298] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
522. 29 <=> v1_xboole_0(k1_gate_1(sK10)) [avatar definition]
523. v1_xboole_0(k1_gate_1(sK10)) <- (29) [avatar component clause 522] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
524. ~v1_xboole_0(k1_gate_1(sK10)) <- (~29) [avatar component clause 522] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
525. ~29 | 24 [avatar split clause 520,462,522]
527. v1_xboole_0(sK10) <- (~29) [resolution 524,295] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
529. $false <- (~12, ~29) [subsumption resolution 527,365] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
530. 12 | 29 [avatar contradiction clause 529]
534. v1_xboole_0(k1_gate_1(sK9)) | v1_xboole_0(k1_gate_1(sK8)) <- (25) [resolution 299,469] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
535. v1_xboole_0(k1_gate_1(sK11)) | v1_xboole_0(k1_gate_1(sK10)) <- (24) [resolution 299,463] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
538. v1_xboole_0(k1_gate_1(sK8)) | v1_xboole_0(sK9) <- (18) [resolution 299,401] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
542. v1_xboole_0(k1_gate_1(sK8)) <- (~11, 25) [subsumption resolution 534,360] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
544. 30 <=> v1_xboole_0(k1_gate_1(sK8)) [avatar definition]
545. ~v1_xboole_0(k1_gate_1(sK8)) <- (~30) [avatar component clause 544] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
546. v1_xboole_0(k1_gate_1(sK8)) <- (30) [avatar component clause 544] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
547. 30 | 11 | ~25 [avatar split clause 542,467,359,544]
548. v1_xboole_0(k1_gate_1(sK10)) <- (~22, 24) [subsumption resolution 535,443] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
549. 29 | 22 | ~24 [avatar split clause 548,462,441,522]
552. ~v1_xboole_0(sK10) <- (29) [resolution 523,296] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
553. ~v1_xboole_0(k1_gate_1(sK8)) <- (~25) [resolution 468,298] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
554. ~30 | 25 [avatar split clause 553,467,544]
555. v1_xboole_0(sK8) <- (~30) [resolution 545,295] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
557. ~v1_xboole_0(sK11) <- (~28) [resolution 454,483] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
564. $false <- (10, ~28) [subsumption resolution 557,354] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
565. ~10 | 28 [avatar contradiction clause 564]
568. v1_xboole_0(k2_gate_1(sK10,sK11)) <- (13) [forward demodulation 371,297] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
569. 28 | ~13 [avatar split clause 568,369,482]
570. v1_xboole_0(sK11) | v1_xboole_0(sK10) <- (28) [resolution 484,299] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
573. ~v1_xboole_0(sK8) <- (~26) [resolution 473,298] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
577. ~9 | 26 [avatar split clause 573,472,349]
578. 9 | 30 [avatar split clause 555,544,349]
586. v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) <- (19) [forward demodulation 407,297] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
587. 26 | ~19 [avatar split clause 586,405,472]
589. ~v1_xboole_0(sK11) <- (22) [resolution 442,296] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
593. ~10 | ~22 [avatar split clause 589,441,353]
594. 12 | 10 | ~28 [avatar split clause 570,482,353,363]
595. ~12 | ~29 [avatar split clause 552,522,363]
596. 21 | 30 | ~18 [avatar split clause 538,399,544,418]
600. ~v1_xboole_0(k2_gate_1(sK8,sK9)) <- (~17) [forward demodulation 394,297] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
601. ~27 | 17 [avatar split clause 600,393,477]
605. ~v1_xboole_0(sK8) <- (~27) [resolution 479,298] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
610. ~v1_xboole_0(k1_gate_1(sK8)) <- (~18) [resolution 400,454] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
615. ~9 | 27 [avatar split clause 605,477,349]
616. ~30 | 18 [avatar split clause 610,399,544]
623. v1_xboole_0(k1_gate_1(sK10)) | v1_xboole_0(sK11) <- (14) [resolution 377,299] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
628. 10 | 29 | ~14 [avatar split clause 623,375,522,353]
629. v1_xboole_0(k2_gate_1(sK8,sK9)) <- (17) [forward demodulation 395,297] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
630. 27 | ~17 [avatar split clause 629,393,477]
635. v1_xboole_0(sK9) | v1_xboole_0(sK8) <- (27) [resolution 478,299] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
641. v1_xboole_0(k1_gate_1(sK9)) | v1_xboole_0(sK8) <- (26) [resolution 474,299] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
643. ~v1_xboole_0(sK8) <- (30) [resolution 546,296] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
645. $false <- (9, 30) [subsumption resolution 643,351] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
646. ~9 | ~30 [avatar contradiction clause 645]
647. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) <- (~20) [forward demodulation 412,297] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
648. ~25 | 20 [avatar split clause 647,411,467]
653. v1_xboole_0(k1_gate_1(sK11)) | v1_xboole_0(sK10) <- (23) [resolution 458,299] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
655. v1_xboole_0(sK10) <- (~22, 23) [subsumption resolution 653,443] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
656. $false <- (~12, ~22, 23) [subsumption resolution 655,365] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
657. 12 | 22 | ~23 [avatar contradiction clause 656]
658. v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) <- (20) [forward demodulation 413,297] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
659. 25 | ~20 [avatar split clause 658,411,467]
660. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) <- (~16) [forward demodulation 388,297] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
661. ~24 | 16 [avatar split clause 660,387,462]
662. ~v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) <- (~15) [forward demodulation 382,297] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
663. ~23 | 15 [avatar split clause 662,381,457]
670. ~v1_xboole_0(sK10) <- (~23) [resolution 459,298] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
672. $false <- (12, ~23) [subsumption resolution 670,364] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
673. ~12 | 23 [avatar contradiction clause 672]
674. 9 | 21 | ~27 [avatar split clause 635,477,418,349]
675. 9 | 11 | ~26 [avatar split clause 641,472,359,349]
676. v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) <- (16) [forward demodulation 389,297] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
677. 24 | ~16 [avatar split clause 676,387,462]
678. v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) <- (15) [forward demodulation 383,297] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
679. 23 | ~15 [avatar split clause 678,381,457]
680. ~v1_xboole_0(k2_gate_1(sK10,sK11)) <- (~13) [forward demodulation 370,297] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
681. ~28 | 13 [avatar split clause 680,369,482]
682. ~v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) <- (~19) [forward demodulation 406,297] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
683. ~26 | 19 [avatar split clause 682,405,472]
684. v1_xboole_0(sK9) <- (~11) [resolution 360,295] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
685. 21 | 11 [avatar split clause 684,359,418]
688. ~v1_xboole_0(k1_gate_1(sK10)) <- (~14) [resolution 376,454] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
690. ~29 | 14 [avatar split clause 688,375,522]
691. $false [avatar sat refutation 339,343,345,346,356,357,366,367,372,373,378,379,384,385,390,391,396,397,402,403,408,409,414,415,437,439,444,445,446,447,501,508,517,518,525,530,547,549,554,565,569,577,578,587,593,594,595,596,601,615,616,628,630,646,648,657,659,661,663,673,674,675,677,679,681,683,685,690] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:676,thDist:-676}
% SZS output end Proof for gate_3__t1_gate_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 12
% 
% >>> Saturation
% Initial clauses: 281
% Generated clauses: 634
% Active clauses: 136
% Passive clauses: 163
% Final active clauses: 22
% Final passive clauses: 9
% Main loop iterations started: 136
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 19
% Fw demodulations: 29
% Bw demodulations: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 243
% Forward subsumptions: 47
% 
% >>> Generating Inferences
% Binary resolution: 69
% Unit resulting resolution: 13
% Forward superposition: 4
% Backward superposition: 2
% 
% >>> AVATAR
% Split clauses: 42
% Split components: 180
% Sat splitting refutations: 27
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 131
% SAT solver binary clauses: 100
% 
% Memory used [KB]: 5245
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% resolution: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

