i: [1,130,0,2,1,1,0,0] t21_grfunc_1
a: [3,212,0,2,1,130]
s: 212
i: [1,129,0,2,1,1,0,0] t21_grfunc_1
a: [3,217,0,2,1,129]
s: 217
i: [1,131,0,3,1,1,0,0] t21_grfunc_1
a: [3,207,0,3,1,131]
s: 207
r: 217
r: 212
d: [2,329,0,2,1,2,23,217,207]
a: [3,334,0,2,1,329]
s: 334
d: [2,330,0,2,1,2,23,212,207]
a: [3,339,0,2,1,330]
s: 339
i: [1,132,0,4,1,1,0,0] t21_grfunc_1
a: [3,202,0,4,1,132]
s: 202
d: [2,341,0,4,1,2,22,202,207]
a: [3,345,0,4,1,341]
s: 345
i: [1,133,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,222,0,2,1,133]
s: 222
i: [1,182,0,2,1,0,0,255] rc7_funct_1
a: [3,242,0,2,1,182]
s: 242
i: [1,181,0,2,1,0,0,255] rc7_funct_1
a: [3,247,0,2,1,181]
s: 247
i: [1,184,0,2,1,0,0,2] rc1_relat_1
a: [3,252,0,2,1,184]
s: 252
i: [1,183,0,2,1,0,0,2] rc1_relat_1
a: [3,257,0,2,1,183]
s: 257
i: [1,187,0,2,1,0,0,255] rc4_funct_1
a: [3,262,0,2,1,187]
s: 262
i: [1,186,0,2,1,0,0,255] rc4_funct_1
a: [3,267,0,2,1,186]
s: 267
i: [1,185,0,2,1,0,0,255] rc4_funct_1
a: [3,272,0,2,1,185]
s: 272
i: [1,188,0,2,1,0,0,255] rc2_relat_1
a: [3,277,0,2,1,188]
s: 277
i: [1,190,0,2,1,0,0,255] rc3_funct_1
a: [3,282,0,2,1,190]
s: 282
i: [1,189,0,2,1,0,0,255] rc3_funct_1
a: [3,287,0,2,1,189]
s: 287
i: [1,192,0,2,1,0,0,1] rc1_funct_1
a: [3,292,0,2,1,192]
s: 292
i: [1,191,0,2,1,0,0,1] rc1_funct_1
a: [3,297,0,2,1,191]
s: 297
i: [1,195,0,2,1,0,0,255] rc5_funct_1
a: [3,302,0,2,1,195]
s: 302
i: [1,194,0,2,1,0,0,255] rc5_funct_1
a: [3,307,0,2,1,194]
s: 307
i: [1,193,0,2,1,0,0,255] rc5_funct_1
a: [3,312,0,2,1,193]
s: 312
i: [1,198,0,2,1,0,0,2] rc2_funct_1
a: [3,317,0,2,1,198]
s: 317
i: [1,197,0,2,1,0,0,2] rc2_funct_1
a: [3,322,0,2,1,197]
s: 322
i: [1,196,0,2,1,0,0,2] rc2_funct_1
a: [3,327,0,2,1,196]
s: 327
i: [1,134,0,3,1,0,0,255] d2_xboole_0
a: [3,227,0,3,1,134]
s: 227
r: 222
d: [2,347,0,2,1,2,23,222,227]
a: [3,351,0,2,1,347]
s: 351
i: [1,139,0,4,2,0,0,255] cc7_funct_1
s: 139
i: [1,140,0,4,2,0,0,1] cc1_funct_1
s: 140
i: [1,141,0,4,2,0,0,2] cc1_relat_1
s: 141
d: [2,353,1,2,1,1,40,139,351]
a: [3,357,1,2,1,353]
s: 357
i: [1,179,0,4,1,0,0,255] existence_m1_subset_1
s: 179
i: [1,136,0,4,1,0,0,1] t38_relat_1
a: [3,232,0,4,1,136]
s: 232
i: [1,135,0,4,1,0,0,1] t38_relat_1
a: [3,237,0,4,1,135]
s: 237
i: [1,142,0,5,2,0,0,1] t6_boole
s: 142
d: [2,366,1,6,2,1,44,345,142]
s: 366
i: [1,143,0,5,2,0,0,2] fc4_xtuple_0
s: 143
i: [1,144,0,5,2,0,0,2] fc10_relat_1
s: 144
i: [1,145,0,5,2,0,0,2] fc5_xtuple_0
s: 145
i: [1,146,0,5,2,0,0,2] fc11_relat_1
s: 146
i: [1,178,0,5,2,0,0,255] t7_boole
s: 178
i: [1,175,0,6,2,0,0,255] t1_subset
s: 175
d: [2,369,1,5,2,0,40,143,141]
s: 369
i: [1,180,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 180
i: [1,137,0,7,3,0,0,255] cc8_funct_1
s: 137
d: [2,370,1,5,2,0,40,143,140]
s: 370
i: [1,138,0,7,3,0,0,255] cc8_funct_1
s: 138
d: [2,371,1,5,2,0,40,143,139]
s: 371
i: [1,148,0,7,3,0,0,2] fc8_relat_1
s: 148
i: [1,149,0,7,3,0,0,2] fc9_relat_1
s: 149
d: [2,373,1,5,2,0,40,145,141]
s: 373
i: [1,153,0,7,3,0,0,2] fc24_relat_1
s: 153
d: [2,374,1,5,2,0,40,145,140]
s: 374
i: [1,154,0,7,3,0,0,3] fc25_relat_1
s: 154
i: [1,171,0,7,3,0,0,1] dt_k2_funct_1
s: 171
d: [2,375,1,5,2,0,40,145,139]
s: 375
i: [1,172,0,7,3,0,0,1] dt_k2_funct_1
s: 172
d: [2,378,1,5,2,0,40,137,179]
s: 378
i: [1,177,0,7,3,0,0,1] t8_boole
s: 177
i: [1,152,0,8,4,0,0,255] cc6_funct_1
s: 152
d: [2,380,1,5,2,0,40,138,179]
s: 380
i: [1,157,0,8,4,0,0,2] cc2_funct_1
s: 157
d: [2,393,0,6,3,0,21,157,140]
d: [2,394,0,4,2,0,21,393,141]
s: 394
i: [1,160,0,8,4,0,0,255] cc4_funct_1
s: 160
d: [2,403,0,6,3,0,21,160,140]
d: [2,404,0,4,2,0,21,403,141]
s: 404
d: [2,395,1,2,1,1,40,394,351]
a: [3,401,1,2,1,395]
s: 401
i: [1,164,0,8,4,0,0,255] cc5_funct_1
s: 164
d: [2,405,1,2,1,1,40,404,351]
a: [3,411,1,2,1,405]
s: 411
i: [1,176,0,8,3,0,0,255] t2_subset
s: 176
i: [1,147,0,9,4,0,0,2] fc18_funct_1
s: 147
d: [2,396,1,5,2,0,40,394,143]
s: 396
i: [1,163,0,9,4,0,0,255] fc13_funct_1
s: 163
d: [2,397,1,5,2,0,40,394,145]
s: 397
i: [1,167,0,9,4,0,0,1] fc6_funct_1
s: 167
i: [1,168,0,9,4,0,0,1] fc6_funct_1
s: 168
d: [2,406,1,5,2,0,40,404,143]
s: 406
i: [1,169,0,9,4,0,0,1] fc6_funct_1
s: 169
d: [2,407,1,5,2,0,40,404,145]
s: 407
i: [1,170,0,9,4,0,0,255] fc12_funct_1
s: 170
i: [1,173,0,12,4,0,0,1] t33_funct_1
s: 173
d: [2,428,1,8,4,1,44,170,232]
d: [2,430,1,6,3,2,21,428,411]
d: [2,431,1,4,2,3,21,430,339]
d: [2,432,1,2,1,4,21,431,334]
a: [3,436,1,2,1,432]
s: 436
i: [1,174,0,12,4,0,0,1] t33_funct_1
s: 174
d: [2,452,2,4,2,1,44,436,142]
s: 452
d: [2,362,1,6,2,1,44,232,142]
s: 362
d: [2,363,1,6,2,1,44,237,142]
s: 363
d: [2,470,3,5,2,1,40,452,143]
s: 470
d: [2,415,1,6,2,0,40,176,179]
s: 415
d: [2,471,3,5,2,1,40,452,145]
s: 471
d: [2,361,1,7,3,0,44,142,142]
s: 361
d: [2,438,1,12,5,0,44,143,173]
s: 438
d: [2,512,4,12,5,1,44,470,173]
s: 512
d: [2,439,1,18,7,0,44,147,173]
d: [2,447,1,15,6,0,21,439,172]
d: [2,448,1,12,5,0,21,447,171]
s: 448
d: [2,527,2,11,5,1,44,448,232]
d: [2,535,2,9,4,2,21,527,436]
d: [2,536,2,7,3,3,21,535,401]
d: [2,537,2,5,2,4,21,536,339]
d: [2,538,2,3,1,5,21,537,334]
a: [3,542,2,3,1,538]
s: 542
d: [2,440,1,15,6,0,44,148,173]
d: [2,449,1,12,5,0,21,440,171]
s: 449
d: [2,441,1,15,6,0,44,153,173]
d: [2,450,1,12,5,0,21,441,171]
s: 450
d: [2,567,2,11,5,1,44,449,232]
d: [2,574,2,9,4,2,21,567,351]
d: [2,575,2,7,3,3,21,574,401]
d: [2,576,2,5,2,4,21,575,339]
d: [2,577,2,3,1,5,21,576,334]
a: [3,581,2,3,1,577]
s: 581
d: [2,594,3,0,0,2,52,345,581,142]
e: 594
d: [2,595,3,0,0,3,52,351,345,581,177]
e: 595
d: [2,596,3,0,0,3,52,351,345,581,177]
e: 596
f: 594,577,395,347,341,330,329,136,134,133,129,130,131,132
% Refutation found. Thanks to Tanya!
% SZS status Theorem for grfunc_1__t21_grfunc_1
% SZS output start Proof for grfunc_1__t21_grfunc_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (k1_xboole_0 = X0 => k1_xboole_0 = k2_funct_1(X0))) [input t21_grfunc_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (k1_xboole_0 = X0 => k1_xboole_0 = k2_funct_1(X0))) [negated conjecture 1]
4. ! [X0] : (v1_xboole_0(X0) => v1_funct_1(X0)) [input cc1_funct_1]
5. ! [X0] : (v1_xboole_0(X0) => v1_relat_1(X0)) [input cc1_relat_1]
6. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0) & v1_xboole_0(X0)) => (v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0))) [input cc2_funct_1]
14. k1_xboole_0 = o_0_0_xboole_0 [input d2_xboole_0]
17. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0)))) [input dt_k2_funct_1]
20. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
33. ! [X0] : ((v1_relat_1(X0) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k9_xtuple_0(X0))) [input fc8_relat_1]
45. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))))) [input t33_funct_1]
46. k1_xboole_0 = k10_xtuple_0(k1_xboole_0) & k1_xboole_0 = k9_xtuple_0(k1_xboole_0) [input t38_relat_1]
47. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
60. ? [X0] : ((k1_xboole_0 != k2_funct_1(X0) & k1_xboole_0 = X0) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
61. ? [X0] : (k1_xboole_0 != k2_funct_1(X0) & k1_xboole_0 = X0 & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 60]
64. ! [X0] : (v1_funct_1(X0) | ~v1_xboole_0(X0)) [ennf transformation 4]
65. ! [X0] : (v1_relat_1(X0) | ~v1_xboole_0(X0)) [ennf transformation 5]
66. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 47]
73. ! [X0] : (~v1_xboole_0(k9_xtuple_0(X0)) | (~v1_relat_1(X0) | v1_xboole_0(X0))) [ennf transformation 33]
74. ! [X0] : (~v1_xboole_0(k9_xtuple_0(X0)) | ~v1_relat_1(X0) | v1_xboole_0(X0)) [flattening 73]
83. ! [X0] : ((v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_xboole_0(X0))) [ennf transformation 6]
84. ! [X0] : ((v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_xboole_0(X0)) [flattening 83]
99. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 17]
100. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 99]
101. ! [X0] : (((k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 45]
102. ! [X0] : ((k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 101]
109. ? [X0] : (k1_xboole_0 != k2_funct_1(X0) & k1_xboole_0 = X0 & v1_funct_1(X0) & v1_relat_1(X0)) => (k1_xboole_0 != k2_funct_1(sK0) & k1_xboole_0 = sK0 & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
110. k1_xboole_0 != k2_funct_1(sK0) & k1_xboole_0 = sK0 & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 61,109]
129. v1_relat_1(sK0) [cnf transformation 110] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
130. v1_funct_1(sK0) [cnf transformation 110] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
131. k1_xboole_0 = sK0 [cnf transformation 110] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
132. k1_xboole_0 != k2_funct_1(sK0) [cnf transformation 110] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
133. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 20] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
134. k1_xboole_0 = o_0_0_xboole_0 [cnf transformation 14] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
136. k1_xboole_0 = k10_xtuple_0(k1_xboole_0) [cnf transformation 46] {a:0,w:4,wCS:1600,sine:1,thAx:0,allAx:1,thDist:-1}
140. ~v1_xboole_0(X0) | v1_funct_1(X0) [cnf transformation 64] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
141. ~v1_xboole_0(X0) | v1_relat_1(X0) [cnf transformation 65] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
142. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 66] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
148. ~v1_xboole_0(k9_xtuple_0(X0)) | ~v1_relat_1(X0) | v1_xboole_0(X0) [cnf transformation 74] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
157. v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_xboole_0(X0) [cnf transformation 84] {a:0,w:8,wCS:3200,sine:2,thAx:0,allAx:1,thDist:-1}
171. v1_relat_1(k2_funct_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 100] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
173. k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 102] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
200. 1 <=> k1_xboole_0 = k2_funct_1(sK0) [avatar definition]
202. k1_xboole_0 != k2_funct_1(sK0) <- (~1) [avatar component clause 200] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
203. ~1 [avatar split clause 132,200]
205. 2 <=> k1_xboole_0 = sK0 [avatar definition]
207. k1_xboole_0 = sK0 <- (2) [avatar component clause 205] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
208. 2 [avatar split clause 131,205]
210. 3 <=> v1_funct_1(sK0) [avatar definition]
212. v1_funct_1(sK0) <- (3) [avatar component clause 210] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
213. 3 [avatar split clause 130,210]
215. 4 <=> v1_relat_1(sK0) [avatar definition]
217. v1_relat_1(sK0) <- (4) [avatar component clause 215] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
218. 4 [avatar split clause 129,215]
220. 5 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
222. v1_xboole_0(o_0_0_xboole_0) <- (5) [avatar component clause 220] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
223. 5 [avatar split clause 133,220]
225. 6 <=> k1_xboole_0 = o_0_0_xboole_0 [avatar definition]
227. k1_xboole_0 = o_0_0_xboole_0 <- (6) [avatar component clause 225] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
228. 6 [avatar split clause 134,225]
230. 7 <=> k1_xboole_0 = k10_xtuple_0(k1_xboole_0) [avatar definition]
232. k1_xboole_0 = k10_xtuple_0(k1_xboole_0) <- (7) [avatar component clause 230] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
233. 7 [avatar split clause 136,230]
329. v1_relat_1(k1_xboole_0) <- (2, 4) [backward demodulation 217,207] {a:0,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
330. v1_funct_1(k1_xboole_0) <- (2, 3) [backward demodulation 212,207] {a:0,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
332. 27 <=> v1_relat_1(k1_xboole_0) [avatar definition]
334. v1_relat_1(k1_xboole_0) <- (27) [avatar component clause 332] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
335. 27 | ~2 | ~4 [avatar split clause 329,215,205,332]
337. 28 <=> v1_funct_1(k1_xboole_0) [avatar definition]
339. v1_funct_1(k1_xboole_0) <- (28) [avatar component clause 337] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
340. 28 | ~2 | ~3 [avatar split clause 330,210,205,337]
341. k1_xboole_0 != k2_funct_1(k1_xboole_0) <- (~1, 2) [forward demodulation 202,207] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
343. 29 <=> k1_xboole_0 = k2_funct_1(k1_xboole_0) [avatar definition]
345. k1_xboole_0 != k2_funct_1(k1_xboole_0) <- (~29) [avatar component clause 343] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
346. ~29 | 1 | ~2 [avatar split clause 341,205,200,343]
347. v1_xboole_0(k1_xboole_0) <- (5, 6) [backward demodulation 222,227] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
349. 30 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
351. v1_xboole_0(k1_xboole_0) <- (30) [avatar component clause 349] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
352. 30 | ~5 | ~6 [avatar split clause 347,225,220,349]
393. v2_funct_1(X0) | ~v1_relat_1(X0) | ~v1_xboole_0(X0) [subsumption resolution 157,140] {a:0,w:6,wCS:2400,sine:1,thAx:0,allAx:1,thDist:-1}
394. ~v1_xboole_0(X0) | v2_funct_1(X0) [subsumption resolution 393,141] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
395. v2_funct_1(k1_xboole_0) <- (30) [resolution 394,351] {a:1,w:2,wCS:800,sine:1,thAx:0,allAx:2,thDist:-2}
399. 32 <=> v2_funct_1(k1_xboole_0) [avatar definition]
401. v2_funct_1(k1_xboole_0) <- (32) [avatar component clause 399] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
402. 32 | ~30 [avatar split clause 395,349,399]
440. ~v1_xboole_0(k10_xtuple_0(X2)) | ~v1_relat_1(k2_funct_1(X2)) | v1_xboole_0(k2_funct_1(X2)) | ~v2_funct_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [superposition 148,173] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
449. ~v1_xboole_0(k10_xtuple_0(X2)) | v1_xboole_0(k2_funct_1(X2)) | ~v2_funct_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [subsumption resolution 440,171] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
567. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(k2_funct_1(k1_xboole_0)) | ~v2_funct_1(k1_xboole_0) | ~v1_funct_1(k1_xboole_0) | ~v1_relat_1(k1_xboole_0) <- (7) [superposition 449,232] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:3,thDist:-3}
574. v1_xboole_0(k2_funct_1(k1_xboole_0)) | ~v2_funct_1(k1_xboole_0) | ~v1_funct_1(k1_xboole_0) | ~v1_relat_1(k1_xboole_0) <- (7, 30) [subsumption resolution 567,351] {a:2,w:9,wCS:3600,sine:1,thAx:0,allAx:3,thDist:-3}
575. v1_xboole_0(k2_funct_1(k1_xboole_0)) | ~v1_funct_1(k1_xboole_0) | ~v1_relat_1(k1_xboole_0) <- (7, 30, 32) [subsumption resolution 574,401] {a:2,w:7,wCS:2800,sine:1,thAx:0,allAx:3,thDist:-3}
576. v1_xboole_0(k2_funct_1(k1_xboole_0)) | ~v1_relat_1(k1_xboole_0) <- (7, 28, 30, 32) [subsumption resolution 575,339] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
577. v1_xboole_0(k2_funct_1(k1_xboole_0)) <- (7, 27, 28, 30, 32) [subsumption resolution 576,334] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
579. 39 <=> v1_xboole_0(k2_funct_1(k1_xboole_0)) [avatar definition]
581. v1_xboole_0(k2_funct_1(k1_xboole_0)) <- (39) [avatar component clause 579] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
582. 39 | ~7 | ~27 | ~28 | ~30 | ~32 [avatar split clause 577,399,349,337,332,230,579]
594. $false <- (~29, 39) [unit resulting resolution 345,581,142] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
606. 29 | ~39 [avatar contradiction clause 594]
617. $false [avatar sat refutation 203,208,213,218,223,228,233,335,340,346,352,402,582,606] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:158,thDist:-158}
% SZS output end Proof for grfunc_1__t21_grfunc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 48
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 3
% 
% >>> Saturation
% Initial clauses: 70
% Generated clauses: 329
% Active clauses: 88
% Passive clauses: 131
% Final active clauses: 85
% Final passive clauses: 34
% Main loop iterations started: 97
% 
% >>> Simplifying Inferences
% Duplicate literals: 24
% Fw subsumption resolutions: 53
% Fw demodulations: 4
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 30
% Forward subsumptions: 52
% 
% >>> Generating Inferences
% Binary resolution: 51
% Unit resulting resolution: 3
% Forward superposition: 38
% Backward superposition: 47
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 3
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 42
% SAT solver unit clauses: 26
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

