i: [1,45,0,9,2,1,0,0] t57_zfmisc_1
a: [3,69,0,3,1,45]
s: 69
a: [3,66,0,6,1,45]
s: 66
i: [1,46,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,76,0,2,1,46]
s: 76
i: [1,47,0,2,1,0,0,2] fc1_xboole_0
a: [3,81,0,2,1,47]
s: 81
i: [1,61,0,2,1,0,0,2] rc2_xboole_0
a: [3,91,0,2,1,61]
s: 91
i: [1,62,0,2,1,0,0,2] rc1_xboole_0
a: [3,96,0,2,1,62]
s: 96
i: [1,49,0,3,1,0,0,1] fc2_xboole_0
s: 49
i: [1,48,0,3,1,0,0,255] d2_xboole_0
a: [3,86,0,3,1,48]
s: 86
r: 76
i: [1,50,0,5,1,0,0,1] t4_boole
s: 50
i: [1,51,0,5,1,0,0,1] t3_boole
s: 51
i: [1,52,0,5,2,0,0,2] t6_boole
s: 52
i: [1,57,0,5,2,0,0,1] t7_boole
s: 57
i: [1,59,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 59
i: [1,60,0,6,2,0,0,2] symmetry_r1_xboole_0
s: 60
i: [1,53,0,7,2,0,0,1] l20_zfmisc_1
s: 53
i: [1,56,0,7,3,0,0,1] t8_boole
s: 56
i: [1,58,0,7,2,0,0,1] l18_zfmisc_1
s: 58
i: [1,54,0,8,2,0,0,1] t83_xboole_1
s: 54
d: [2,105,1,7,2,1,44,66,54]
d: [2,106,1,4,1,1,19,105]
a: [3,110,1,4,1,106]
s: 110
i: [1,55,0,8,2,0,0,1] t83_xboole_1
s: 55
d: [2,113,1,6,2,0,44,55,51]
d: [2,115,1,3,1,0,19,113]
s: 115
d: [2,112,1,6,2,0,44,55,50]
d: [2,116,1,3,1,0,19,112]
s: 116
d: [2,99,1,7,3,0,44,52,52]
s: 99
d: [2,117,2,3,1,0,40,115,60]
s: 117
d: [2,100,1,7,2,0,44,51,52]
s: 100
d: [2,101,1,7,2,0,44,50,52]
s: 101
d: [2,126,2,5,2,1,44,66,101]
d: [2,129,2,2,1,1,19,126]
a: [3,133,2,2,1,129]
s: 133
d: [2,103,1,7,2,0,40,53,60]
s: 103
d: [2,135,2,0,0,2,52,69,110,103]
e: 135
d: [2,136,2,3,1,1,40,103,110]
d: [2,140,2,0,0,2,21,136,69]
e: 140
r: 66
r: 69
a: [3,70,0,3,1,45]
s: 70
a: [3,65,0,6,1,45]
s: 65
d: [2,150,2,0,0,2,52,110,65,55]
e: 150
d: [2,151,1,7,2,1,44,55,65]
d: [2,152,1,4,1,1,19,151]
d: [2,154,1,0,0,2,21,152,110]
e: 154
r: 110
d: [2,143,1,3,1,1,40,70,59]
a: [3,148,1,3,1,143]
s: 148
a: [3,109,1,4,1,106]
s: 109
d: [2,157,2,4,1,1,40,109,60]
a: [3,161,2,4,1,157]
s: 161
d: [2,163,3,0,0,2,52,70,161,58]
e: 163
d: [2,164,3,3,1,1,40,161,58]
d: [2,167,3,0,0,2,21,164,70]
e: 167
i: [1,44,0,9,2,1,0,0] t57_zfmisc_1
f: 167,157,152,136,106,44,45
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t57_zfmisc_1
% SZS output start Proof for zfmisc_1__t57_zfmisc_1
1. ! [X0] : ! [X1] : (k4_xboole_0(X0,k1_tarski(X1)) = X0 <=> ~r2_hidden(X1,X0)) [input t57_zfmisc_1]
2. ~! [X0] : ! [X1] : (k4_xboole_0(X0,k1_tarski(X1)) = X0 <=> ~r2_hidden(X1,X0)) [negated conjecture 1]
11. ! [X0] : ! [X1] : ~(r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [input l18_zfmisc_1]
12. ! [X0] : ! [X1] : (~r2_hidden(X0,X1) => r1_xboole_0(k1_tarski(X0),X1)) [input l20_zfmisc_1]
15. ! [X0,X1] : (r1_xboole_0(X0,X1) => r1_xboole_0(X1,X0)) [input symmetry_r1_xboole_0]
20. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [input t83_xboole_1]
22. ~! [X0,X1] : (k4_xboole_0(X0,k1_tarski(X1)) = X0 <=> ~r2_hidden(X1,X0)) [flattening 2]
23. ! [X0,X1] : (~r2_hidden(X0,X1) => r1_xboole_0(k1_tarski(X0),X1)) [flattening 12]
24. ! [X0,X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [flattening 20]
27. ! [X0,X1] : ~(r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [flattening 11]
28. ? [X0,X1] : (k4_xboole_0(X0,k1_tarski(X1)) = X0 <~> ~r2_hidden(X1,X0)) [ennf transformation 22]
30. ! [X0,X1] : (r1_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1)) [ennf transformation 23]
33. ! [X0,X1] : (~r2_hidden(X0,X1) | ~r1_xboole_0(k1_tarski(X0),X1)) [ennf transformation 27]
35. ! [X0,X1] : (r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 15]
36. ? [X0,X1] : ((r2_hidden(X1,X0) | k4_xboole_0(X0,k1_tarski(X1)) != X0) & (~r2_hidden(X1,X0) | k4_xboole_0(X0,k1_tarski(X1)) = X0)) [nnf transformation 28]
37. ? [X0,X1] : ((r2_hidden(X1,X0) | k4_xboole_0(X0,k1_tarski(X1)) != X0) & (~r2_hidden(X1,X0) | k4_xboole_0(X0,k1_tarski(X1)) = X0)) => ((r2_hidden(sK1,sK0) | sK0 != k4_xboole_0(sK0,k1_tarski(sK1))) & (~r2_hidden(sK1,sK0) | sK0 = k4_xboole_0(sK0,k1_tarski(sK1)))) [choice axiom]
38. (r2_hidden(sK1,sK0) | sK0 != k4_xboole_0(sK0,k1_tarski(sK1))) & (~r2_hidden(sK1,sK0) | sK0 = k4_xboole_0(sK0,k1_tarski(sK1))) [skolemisation 36,37]
39. ! [X0,X1] : ((r1_xboole_0(X0,X1) | k4_xboole_0(X0,X1) != X0) & (k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1))) [nnf transformation 24]
44. ~r2_hidden(sK1,sK0) | sK0 = k4_xboole_0(sK0,k1_tarski(sK1)) [cnf transformation 38] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. r2_hidden(sK1,sK0) | sK0 != k4_xboole_0(sK0,k1_tarski(sK1)) [cnf transformation 38] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. r1_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1) [cnf transformation 30] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1) [cnf transformation 39] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. k4_xboole_0(X0,X1) != X0 | r1_xboole_0(X0,X1) [cnf transformation 39] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. ~r1_xboole_0(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 33] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. ~r1_xboole_0(X0,X1) | r1_xboole_0(X1,X0) [cnf transformation 35] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
64. 1 <=> sK0 = k4_xboole_0(sK0,k1_tarski(sK1)) [avatar definition]
65. sK0 = k4_xboole_0(sK0,k1_tarski(sK1)) <- (1) [avatar component clause 64] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66. sK0 != k4_xboole_0(sK0,k1_tarski(sK1)) <- (~1) [avatar component clause 64] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. 2 <=> r2_hidden(sK1,sK0) [avatar definition]
70. r2_hidden(sK1,sK0) <- (2) [avatar component clause 68] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. ~1 | 2 [avatar split clause 45,68,64]
72. 1 | ~2 [avatar split clause 44,68,64]
103. r1_xboole_0(X1,k1_tarski(X0)) | r2_hidden(X0,X1) [resolution 53,60] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
105. sK0 != sK0 | ~r1_xboole_0(sK0,k1_tarski(sK1)) <- (~1) [superposition 66,54] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
106. ~r1_xboole_0(sK0,k1_tarski(sK1)) <- (~1) [trivial inequality removal 105] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
108. 8 <=> r1_xboole_0(sK0,k1_tarski(sK1)) [avatar definition]
109. r1_xboole_0(sK0,k1_tarski(sK1)) <- (8) [avatar component clause 108] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. ~r1_xboole_0(sK0,k1_tarski(sK1)) <- (~8) [avatar component clause 108] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
111. ~8 | 1 [avatar split clause 106,64,108]
136. r2_hidden(sK1,sK0) <- (~8) [resolution 103,110] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
142. 2 | 8 [avatar split clause 136,108,68]
151. sK0 != sK0 | r1_xboole_0(sK0,k1_tarski(sK1)) <- (1) [superposition 55,65] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
152. r1_xboole_0(sK0,k1_tarski(sK1)) <- (1) [trivial inequality removal 151] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
156. 8 | ~1 [avatar split clause 152,64,108]
157. r1_xboole_0(k1_tarski(sK1),sK0) <- (8) [resolution 109,60] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
159. 11 <=> r1_xboole_0(k1_tarski(sK1),sK0) [avatar definition]
161. r1_xboole_0(k1_tarski(sK1),sK0) <- (11) [avatar component clause 159] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
162. 11 | ~8 [avatar split clause 157,108,159]
164. ~r2_hidden(sK1,sK0) <- (11) [resolution 161,58] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
167. $false <- (2, 11) [subsumption resolution 164,70] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
168. ~2 | ~11 [avatar contradiction clause 167]
169. $false [avatar sat refutation 71,72,111,142,156,162,168] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:87,thDist:-87}
% SZS output end Proof for zfmisc_1__t57_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 79
% Active clauses: 31
% Passive clauses: 39
% Final active clauses: 27
% Final passive clauses: 6
% Main loop iterations started: 33
% 
% >>> Simplifying Inferences
% Trivial inequalities: 8
% Fw subsumption resolutions: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 3
% Forward superposition: 5
% Backward superposition: 9
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 5
% SAT solver binary clauses: 14
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward superposition index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

