i: [1,59,0,6,2,1,0,0] t110_zfmisc_1
a: [3,88,0,3,1,59]
s: 88
i: [1,58,0,3,1,1,0,0] t110_zfmisc_1
a: [3,97,0,3,1,58]
s: 97
i: [1,57,0,3,1,1,0,0] t110_zfmisc_1
a: [3,102,0,3,1,57]
s: 102
i: [1,56,0,7,1,1,0,0] t110_zfmisc_1
a: [3,107,0,7,1,56]
s: 107
i: [1,60,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,112,0,2,1,60]
s: 112
i: [1,61,0,2,1,0,0,1] fc1_xboole_0
a: [3,117,0,2,1,61]
s: 117
i: [1,81,0,2,1,0,0,2] rc2_xboole_0
a: [3,127,0,2,1,81]
s: 127
i: [1,82,0,2,1,0,0,2] rc1_xboole_0
a: [3,132,0,2,1,82]
s: 132
i: [1,63,0,3,1,0,0,255] fc2_xboole_0
s: 63
i: [1,62,0,3,1,0,0,255] d2_xboole_0
a: [3,122,0,3,1,62]
s: 122
r: 112
i: [1,77,0,4,1,0,0,2] fc1_zfmisc_1
s: 77
i: [1,78,0,4,1,0,0,255] fc3_xboole_0
s: 78
i: [1,64,0,5,2,0,0,1] t6_boole
s: 64
i: [1,73,0,5,2,0,0,2] t7_boole
s: 73
i: [1,83,0,5,1,0,0,1] t90_zfmisc_1
s: 83
i: [1,84,0,5,1,0,0,1] t90_zfmisc_1
s: 84
i: [1,80,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 80
i: [1,65,0,7,2,0,0,2] t7_xboole_0
s: 65
i: [1,72,0,7,3,0,0,1] t8_boole
s: 72
i: [1,79,0,7,1,0,0,255] commutativity_k2_tarski
s: 79
i: [1,66,0,10,1,0,0,255] d5_tarski
s: 66
d: [2,135,1,7,3,0,44,64,64]
s: 135
i: [1,74,0,10,2,0,0,1] l46_zfmisc_1
s: 74
d: [2,151,1,10,2,1,44,74,107]
s: 151
i: [1,75,0,10,2,0,0,1] l46_zfmisc_1
s: 75
i: [1,69,0,11,3,0,0,1] t90_zfmisc_1
s: 69
d: [2,157,1,11,3,1,44,69,107]
a: [3,171,1,5,1,157]
s: 171
i: [1,67,0,13,3,0,0,2] t1_tarski
s: 67
d: [2,154,1,10,2,1,44,75,107]
s: 154
i: [1,68,0,13,3,0,0,2] t1_tarski
s: 68
i: [1,76,0,13,3,0,0,1] l46_zfmisc_1
s: 76
d: [2,180,2,9,3,1,40,76,154]
a: [3,189,2,3,1,180]
s: 189
d: [2,201,3,0,0,2,52,102,189,65]
e: 201
d: [2,202,3,3,1,1,40,189,65]
d: [2,205,3,0,0,2,21,202,102]
e: 205
r: 189
d: [2,184,1,13,3,1,44,76,107]
s: 184
d: [2,209,2,9,3,1,40,184,75]
a: [3,215,2,3,1,209]
s: 215
r: 184
a: [3,167,1,3,1,157]
s: 167
r: 215
r: 151
r: 107
r: 88
d: [2,237,0,7,1,2,23,107,167]
d: [2,239,0,5,1,2,22,237,84]
d: [2,240,0,0,0,3,21,239,171]
e: 240
r: 167
s: 88
d: [2,212,2,10,3,1,40,184,73]
a: [3,230,2,4,1,212]
s: 230
s: 107
a: [3,192,2,6,2,180]
s: 192
s: 151
d: [2,245,1,10,2,1,44,74,107]
s: 245
d: [2,246,3,7,2,1,40,192,65]
d: [2,248,3,4,1,2,21,246,97]
a: [3,252,3,4,1,248]
s: 252
s: 184
d: [2,257,4,2,1,1,40,252,73]
a: [3,266,4,2,1,257]
s: 266
d: [2,242,1,13,3,1,44,76,107]
s: 242
d: [2,137,1,7,2,0,44,83,64]
s: 137
d: [2,283,2,13,4,1,44,184,137]
d: [2,290,2,8,3,1,21,283,73]
a: [3,294,2,2,1,290]
s: 294
d: [2,138,1,7,2,0,44,84,64]
s: 138
d: [2,299,2,7,2,1,44,138,107]
a: [3,316,2,2,1,299]
s: 316
d: [2,139,1,7,2,0,40,65,80]
s: 139
d: [2,152,1,8,2,0,44,74,83]
s: 152
d: [2,306,2,13,4,1,44,184,138]
d: [2,322,2,8,3,1,21,306,73]
a: [3,326,2,2,1,322]
s: 326
d: [2,153,1,8,2,0,44,74,84]
s: 153
d: [2,256,4,4,1,1,40,252,80]
a: [3,261,4,4,1,256]
s: 261
d: [2,155,1,8,2,0,44,75,83]
s: 155
d: [2,156,1,8,2,0,44,75,84]
s: 156
r: 155
d: [2,181,2,9,3,1,40,76,151]
a: [3,196,2,6,2,181]
s: 196
d: [2,143,1,10,1,0,44,66,79]
s: 143
d: [2,333,3,7,2,1,40,196,65]
d: [2,335,3,4,1,2,21,333,102]
a: [3,339,3,4,1,335]
s: 339
d: [2,144,1,10,1,0,44,66,79]
s: 144
d: [2,146,1,10,1,0,44,66,79]
s: 146
d: [2,351,4,4,1,1,40,339,80]
a: [3,356,4,4,1,351]
s: 356
d: [2,147,1,10,1,0,44,66,79]
s: 147
a: [3,218,2,6,2,209]
s: 218
a: [3,162,1,3,1,157]
s: 162
d: [2,148,1,10,1,0,44,79,66]
s: 148
d: [2,367,3,7,2,1,40,218,65]
a: [3,377,3,4,1,367]
s: 377
d: [2,149,1,10,1,0,44,79,66]
s: 149
d: [2,366,4,4,1,2,40,218,252]
a: [3,372,4,4,1,366]
s: 372
d: [2,174,1,10,3,0,40,67,73]
s: 174
d: [2,183,1,10,3,0,40,76,73]
s: 183
d: [2,382,4,4,1,2,40,377,192]
a: [3,388,4,4,1,382]
s: 388
d: [2,185,1,11,3,0,44,76,83]
s: 185
d: [2,383,4,4,1,1,40,377,80]
a: [3,393,4,4,1,383]
s: 393
d: [2,186,1,11,3,0,44,76,84]
s: 186
d: [2,173,1,13,3,0,40,67,80]
s: 173
d: [2,396,5,4,1,1,40,372,80]
a: [3,401,5,4,1,396]
s: 401
d: [2,182,1,13,3,0,40,76,80]
s: 182
d: [2,418,5,4,1,1,40,388,80]
a: [3,423,5,4,1,418]
s: 423
d: [2,145,1,15,1,0,44,66,66]
s: 145
d: [2,210,2,9,3,1,40,184,74]
a: [3,222,2,6,2,210]
s: 222
r: 222
a: [3,225,2,3,1,210]
s: 225
r: 184
r: 218
d: [2,485,5,0,0,2,40,225,388]
e: 485
d: [2,484,4,0,0,2,40,225,252]
e: 484
d: [2,480,4,0,0,2,52,252,225]
e: 480
d: [2,481,4,0,0,3,52,377,225,192]
e: 481
d: [2,482,3,0,0,2,52,162,225,65]
e: 482
d: [2,483,3,0,0,3,52,117,162,225,174]
e: 483
d: [2,486,3,3,1,1,40,225,65]
d: [2,495,3,0,0,2,21,486,162]
e: 495
r: 225
s: 184
d: [2,467,3,7,2,1,40,222,65]
a: [3,478,3,4,1,467]
s: 478
r: 196
a: [3,199,2,3,1,181]
s: 199
r: 192
d: [2,524,5,0,0,2,40,199,372]
e: 524
d: [2,523,4,0,0,2,40,199,377]
e: 523
d: [2,520,4,0,0,2,52,377,199]
e: 520
d: [2,521,3,0,0,2,52,97,199,65]
e: 521
d: [2,522,3,0,0,3,52,117,97,199,174]
e: 522
d: [2,525,3,3,1,1,40,199,65]
d: [2,533,3,0,0,2,21,525,97]
e: 533
r: 199
s: 192
d: [2,466,4,4,1,2,40,222,339]
a: [3,473,4,4,1,466]
s: 473
s: 196
d: [2,507,4,4,1,2,40,478,196]
a: [3,513,4,4,1,507]
s: 513
s: 218
s: 222
d: [2,508,4,4,1,1,40,478,80]
a: [3,518,4,4,1,508]
s: 518
d: [2,211,2,13,3,1,40,184,80]
s: 211
d: [2,540,5,4,1,1,40,473,80]
a: [3,545,5,4,1,540]
s: 545
d: [2,272,2,13,3,1,40,184,80]
s: 272
d: [2,455,2,13,3,1,44,182,107]
s: 455
d: [2,553,5,4,1,1,40,513,80]
a: [3,558,5,4,1,553]
s: 558
d: [2,501,2,13,3,1,40,184,80]
s: 501
d: [2,407,3,10,3,1,40,174,192]
s: 407
d: [2,425,2,9,3,0,40,185,156]
a: [3,435,2,6,2,425]
s: 435
d: [2,426,2,9,3,0,40,185,152]
a: [3,439,2,6,2,426]
s: 439
r: 218
r: 222
r: 192
r: 196
r: 435
d: [2,584,4,4,1,2,40,439,478]
s: 584
r: 513
r: 478
d: [2,600,5,0,0,3,52,132,584,73]
e: 600
d: [2,601,5,0,0,2,52,584,584,80]
e: 601
d: [2,602,5,0,0,2,52,584,584,80]
e: 602
r: 584
r: 439
s: 192
s: 478
s: 196
s: 218
s: 513
s: 222
d: [2,408,3,10,3,1,40,174,218]
s: 408
a: [3,442,2,3,1,426]
s: 442
r: 152
r: 156
r: 185
r: 186
s: 435
d: [2,539,3,10,3,1,40,192,174]
s: 539
d: [2,345,2,10,1,0,44,143,79]
s: 345
d: [2,551,3,10,3,1,40,196,174]
s: 551
d: [2,346,2,10,1,0,44,143,79]
s: 346
d: [2,348,2,10,1,0,44,79,143]
s: 348
d: [2,564,3,10,3,1,40,218,174]
s: 564
d: [2,349,2,10,1,0,44,79,143]
s: 349
d: [2,569,3,10,3,1,40,222,174]
s: 569
d: [2,358,2,10,1,0,44,146,79]
s: 358
d: [2,359,2,10,1,0,44,146,79]
s: 359
d: [2,622,3,10,3,1,40,196,174]
s: 622
d: [2,404,2,10,3,0,40,174,80]
s: 404
d: [2,628,3,10,3,1,40,218,174]
s: 628
d: [2,427,2,11,3,0,40,185,80]
s: 427
d: [2,446,2,11,3,0,40,186,80]
s: 446
d: [2,636,3,10,3,1,40,222,174]
s: 636
d: [2,451,2,11,3,0,44,182,84]
s: 451
d: [2,574,4,10,3,1,40,407,80]
s: 574
d: [2,452,2,11,3,0,44,182,83]
s: 452
d: [2,343,2,14,1,0,44,143,66]
s: 343
d: [2,639,4,10,3,1,40,408,80]
s: 639
d: [2,344,2,14,1,0,44,143,143]
s: 344
d: [2,660,4,10,3,1,40,551,80]
s: 660
d: [2,362,2,14,1,0,44,146,146]
s: 362
d: [2,364,2,14,1,0,44,66,146]
s: 364
d: [2,665,4,10,3,1,40,569,80]
s: 665
d: [2,347,2,15,1,0,44,66,143]
s: 347
d: [2,247,3,13,3,1,40,192,67]
s: 247
d: [2,360,2,15,1,0,44,146,66]
s: 360
d: [2,361,2,15,1,0,44,146,143]
s: 361
d: [2,716,4,8,2,1,42,247]
a: [3,720,4,5,1,716]
s: 720
d: [2,363,2,15,1,0,44,143,146]
s: 363
d: [2,770,5,5,1,2,40,720,218]
a: [3,776,5,5,1,770]
s: 776
a: [3,91,0,3,1,59]
s: 91
r: 107
r: 154
r: 184
r: 192
r: 211
r: 218
r: 247
r: 252
r: 261
r: 266
r: 377
r: 388
r: 393
r: 407
r: 408
r: 423
r: 574
r: 639
r: 720
r: 776
d: [2,794,0,7,1,2,23,107,91]
a: [3,818,0,7,1,794]
s: 818
d: [2,812,4,5,1,2,23,720,91]
a: [3,825,4,5,1,812]
s: 825
d: [2,796,1,13,3,2,23,184,91]
s: 796
d: [2,771,5,5,1,1,40,720,80]
a: [3,781,5,5,1,771]
s: 781
d: [2,833,1,13,3,1,44,76,818]
s: 833
d: [2,798,2,13,3,2,23,211,91]
s: 798
d: [2,839,5,5,1,1,40,825,80]
a: [3,844,5,5,1,839]
s: 844
d: [2,836,2,13,3,1,44,182,818]
s: 836
d: [2,785,6,5,1,1,40,776,80]
a: [3,792,6,5,1,785]
s: 792
d: [2,849,2,13,3,2,40,796,80]
s: 849
d: [2,456,2,15,1,0,44,145,79]
s: 456
d: [2,368,3,13,3,1,40,218,67]
s: 368
d: [2,457,2,15,1,0,44,145,79]
s: 457
d: [2,538,3,13,3,1,40,192,67]
s: 538
d: [2,461,2,19,1,0,44,145,146]
s: 461
d: [2,458,2,20,1,0,44,145,66]
s: 458
d: [2,550,3,13,3,1,40,196,67]
s: 550
d: [2,459,2,20,1,0,44,145,143]
s: 459
d: [2,921,4,8,2,1,42,550]
d: [2,922,4,5,1,2,21,921,88]
a: [3,926,4,5,1,922]
s: 926
d: [2,462,2,20,1,0,44,66,145]
s: 462
d: [2,465,2,20,1,0,44,146,145]
s: 465
d: [2,959,5,5,1,2,40,926,222]
a: [3,965,5,5,1,959]
s: 965
d: [2,1001,6,0,0,3,52,88,926,965,68]
e: 1001
d: [2,1002,6,8,2,1,40,965,68]
d: [2,1007,6,5,1,2,21,1002,88]
d: [2,1008,6,0,0,3,21,1007,926]
e: 1008
d: [2,232,3,3,1,1,40,215,65]
d: [2,254,2,9,3,1,40,151,76]
d: [2,499,2,9,3,1,40,184,75]
d: [2,500,2,9,3,1,40,184,74]
d: [2,783,6,8,2,1,40,776,68]
d: [2,787,6,3,1,2,21,783,720]
f: 1008,959,922,787,770,716,533,500,499,480,254,248,232,205,180,56,57,58,59
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t110_zfmisc_1
% SZS output start Proof for zfmisc_1__t110_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X2,X3) => ((X1 = X3 & X0 = X2) | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t110_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X2,X3) => ((X1 = X3 & X0 = X2) | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [negated conjecture 1]
17. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [input l46_zfmisc_1]
20. ! [X0] : ! [X1] : (! [X2] : (r2_hidden(X2,X0) <=> r2_hidden(X2,X1)) => X0 = X1) [input t1_tarski]
23. ! [X0] : ~(! [X1] : ~r2_hidden(X1,X0) & k1_xboole_0 != X0) [input t7_xboole_0]
26. ~! [X0,X1,X2,X3] : (k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X2,X3) => ((X1 = X3 & X0 = X2) | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 2]
28. ! [X0,X1] : (! [X2] : (r2_hidden(X2,X0) <=> r2_hidden(X2,X1)) => X0 = X1) [flattening 20]
32. ! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [flattening 17]
33. ? [X0,X1,X2,X3] : (((X1 != X3 | X0 != X2) & k1_xboole_0 != X1 & k1_xboole_0 != X0) & k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X2,X3)) [ennf transformation 26]
34. ? [X0,X1,X2,X3] : ((X1 != X3 | X0 != X2) & k1_xboole_0 != X1 & k1_xboole_0 != X0 & k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X2,X3)) [flattening 33]
36. ! [X0] : (? [X1] : r2_hidden(X1,X0) | k1_xboole_0 = X0) [ennf transformation 23]
37. ! [X0,X1] : (X0 = X1 | ? [X2] : (r2_hidden(X2,X0) <~> r2_hidden(X2,X1))) [ennf transformation 28]
41. ? [X0,X1,X2,X3] : ((X1 != X3 | X0 != X2) & k1_xboole_0 != X1 & k1_xboole_0 != X0 & k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X2,X3)) => ((sK1 != sK3 | sK0 != sK2) & k1_xboole_0 != sK1 & k1_xboole_0 != sK0 & k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK2,sK3)) [choice axiom]
42. (sK1 != sK3 | sK0 != sK2) & k1_xboole_0 != sK1 & k1_xboole_0 != sK0 & k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK2,sK3) [skolemisation 34,41]
43. ! [X0] : (? [X1] : r2_hidden(X1,X0) => r2_hidden(sK4(X0),X0)) [choice axiom]
44. ! [X0] : (r2_hidden(sK4(X0),X0) | k1_xboole_0 = X0) [skolemisation 36,43]
45. ! [X0,X1] : (X0 = X1 | ? [X2] : ((~r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) & (r2_hidden(X2,X1) | r2_hidden(X2,X0)))) [nnf transformation 37]
46. ! [X1,X0] : (? [X2] : ((~r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) & (r2_hidden(X2,X1) | r2_hidden(X2,X0))) => ((~r2_hidden(sK5(X0,X1),X1) | ~r2_hidden(sK5(X0,X1),X0)) & (r2_hidden(sK5(X0,X1),X1) | r2_hidden(sK5(X0,X1),X0)))) [choice axiom]
47. ! [X0,X1] : (X0 = X1 | ((~r2_hidden(sK5(X0,X1),X1) | ~r2_hidden(sK5(X0,X1),X0)) & (r2_hidden(sK5(X0,X1),X1) | r2_hidden(sK5(X0,X1),X0)))) [skolemisation 45,46]
50. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | (~r2_hidden(X1,X3) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [nnf transformation 32]
51. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [flattening 50]
56. k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK2,sK3) [cnf transformation 42] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
57. k1_xboole_0 != sK0 [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. k1_xboole_0 != sK1 [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. sK1 != sK3 | sK0 != sK2 [cnf transformation 42] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. r2_hidden(sK4(X0),X0) | k1_xboole_0 = X0 [cnf transformation 44] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
67. r2_hidden(sK5(X0,X1),X1) | X0 = X1 | r2_hidden(sK5(X0,X1),X0) [cnf transformation 47] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
68. ~r2_hidden(sK5(X0,X1),X1) | X0 = X1 | ~r2_hidden(sK5(X0,X1),X0) [cnf transformation 47] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
74. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X0,X2) [cnf transformation 51] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
75. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X1,X3) [cnf transformation 51] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
76. r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2) [cnf transformation 51] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
86. 1 <=> sK0 = sK2 [avatar definition]
88. sK0 != sK2 <- (~1) [avatar component clause 86] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
90. 2 <=> sK1 = sK3 [avatar definition]
93. ~1 | ~2 [avatar split clause 59,90,86]
95. 3 <=> k1_xboole_0 = sK1 [avatar definition]
97. k1_xboole_0 != sK1 <- (~3) [avatar component clause 95] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
98. ~3 [avatar split clause 58,95]
100. 4 <=> k1_xboole_0 = sK0 [avatar definition]
102. k1_xboole_0 != sK0 <- (~4) [avatar component clause 100] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
103. ~4 [avatar split clause 57,100]
105. 5 <=> k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK2,sK3) [avatar definition]
107. k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK2,sK3) <- (5) [avatar component clause 105] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
108. 5 [avatar split clause 56,105]
151. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(sK0,sK1)) | r2_hidden(X0,sK2) <- (5) [superposition 74,107] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
154. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(sK0,sK1)) | r2_hidden(X1,sK3) <- (5) [superposition 75,107] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
165. 12 <=> k1_xboole_0 = sK2 [avatar definition]
180. ~r2_hidden(X8,sK1) | ~r2_hidden(X9,sK0) | r2_hidden(X8,sK3) <- (5) [resolution 76,154] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
184. r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(sK0,sK1)) | ~r2_hidden(X1,sK3) | ~r2_hidden(X0,sK2) <- (5) [superposition 76,107] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
188. 14 <=> ! [X9] : ~r2_hidden(X9,sK0) [avatar definition]
189. ~r2_hidden(X9,sK0) <- (14) [avatar component clause 188] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
191. 15 <=> ! [X8] : (~r2_hidden(X8,sK1) | r2_hidden(X8,sK3)) [avatar definition]
192. ~r2_hidden(X8,sK1) | r2_hidden(X8,sK3) <- (15) [avatar component clause 191] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
193. 14 | 15 | ~5 [avatar split clause 180,105,191,188]
195. 16 <=> ! [X11] : (~r2_hidden(X11,sK0) | r2_hidden(X11,sK2)) [avatar definition]
196. ~r2_hidden(X11,sK0) | r2_hidden(X11,sK2) <- (16) [avatar component clause 195] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
198. 17 <=> ! [X10] : ~r2_hidden(X10,sK1) [avatar definition]
199. ~r2_hidden(X10,sK1) <- (17) [avatar component clause 198] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
202. k1_xboole_0 = sK0 <- (14) [resolution 189,65] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
205. $false <- (~4, 14) [subsumption resolution 202,102] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
206. 4 | ~14 [avatar contradiction clause 205]
214. 18 <=> ! [X5] : ~r2_hidden(X5,sK2) [avatar definition]
215. ~r2_hidden(X5,sK2) <- (18) [avatar component clause 214] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
217. 19 <=> ! [X4] : (~r2_hidden(X4,sK3) | r2_hidden(X4,sK1)) [avatar definition]
218. ~r2_hidden(X4,sK3) | r2_hidden(X4,sK1) <- (19) [avatar component clause 217] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
221. 20 <=> ! [X7] : (~r2_hidden(X7,sK2) | r2_hidden(X7,sK0)) [avatar definition]
222. ~r2_hidden(X7,sK2) | r2_hidden(X7,sK0) <- (20) [avatar component clause 221] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
224. 21 <=> ! [X6] : ~r2_hidden(X6,sK3) [avatar definition]
225. ~r2_hidden(X6,sK3) <- (21) [avatar component clause 224] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
232. k1_xboole_0 = sK2 <- (18) [resolution 215,65] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
234. 12 | ~18 [avatar split clause 232,214,165]
246. r2_hidden(sK4(sK1),sK3) | k1_xboole_0 = sK1 <- (15) [resolution 192,65] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
247. r2_hidden(sK5(X0,sK1),sK3) | sK1 = X0 | r2_hidden(sK5(X0,sK1),X0) <- (15) [resolution 192,67] {a:3,w:13,wCS:1300,nSel:1,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
248. r2_hidden(sK4(sK1),sK3) <- (~3, 15) [subsumption resolution 246,97] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
250. 23 <=> r2_hidden(sK4(sK1),sK3) [avatar definition]
252. r2_hidden(sK4(sK1),sK3) <- (23) [avatar component clause 250] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
253. 23 | 3 | ~15 [avatar split clause 248,191,95,250]
254. r2_hidden(X0,sK2) | ~r2_hidden(X1,sK1) | ~r2_hidden(X0,sK0) <- (5) [resolution 151,76] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
255. 17 | 16 | ~5 [avatar split clause 254,105,195,198]
480. $false <- (21, 23) [unit resulting resolution 252,225] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
491. ~21 | ~23 [avatar contradiction clause 480]
499. ~r2_hidden(X4,sK3) | ~r2_hidden(X5,sK2) | r2_hidden(X4,sK1) <- (5) [resolution 184,75] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
500. ~r2_hidden(X6,sK3) | ~r2_hidden(X7,sK2) | r2_hidden(X7,sK0) <- (5) [resolution 184,74] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
505. 18 | 19 | ~5 [avatar split clause 499,105,217,214]
506. 20 | 21 | ~5 [avatar split clause 500,105,224,221]
525. k1_xboole_0 = sK1 <- (17) [resolution 199,65] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
533. $false <- (~3, 17) [subsumption resolution 525,97] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
534. 3 | ~17 [avatar contradiction clause 533]
550. r2_hidden(sK5(X0,sK0),sK2) | sK0 = X0 | r2_hidden(sK5(X0,sK0),X0) <- (16) [resolution 196,67] {a:3,w:13,wCS:1300,nSel:1,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
716. r2_hidden(sK5(sK3,sK1),sK3) | sK1 = sK3 <- (15) [factoring 247] {a:4,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
718. 48 <=> r2_hidden(sK5(sK3,sK1),sK3) [avatar definition]
720. r2_hidden(sK5(sK3,sK1),sK3) <- (48) [avatar component clause 718] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
721. 2 | 48 | ~15 [avatar split clause 716,191,718,90]
770. r2_hidden(sK5(sK3,sK1),sK1) <- (19, 48) [resolution 720,218] {a:5,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
774. 49 <=> r2_hidden(sK5(sK3,sK1),sK1) [avatar definition]
776. r2_hidden(sK5(sK3,sK1),sK1) <- (49) [avatar component clause 774] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
777. 49 | ~19 | ~48 [avatar split clause 770,718,217,774]
783. sK1 = sK3 | ~r2_hidden(sK5(sK3,sK1),sK3) <- (49) [resolution 776,68] {a:6,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
787. sK1 = sK3 <- (48, 49) [subsumption resolution 783,720] {a:6,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
788. 2 | ~48 | ~49 [avatar split clause 787,774,718,90]
921. r2_hidden(sK5(sK2,sK0),sK2) | sK0 = sK2 <- (16) [factoring 550] {a:4,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
922. r2_hidden(sK5(sK2,sK0),sK2) <- (~1, 16) [subsumption resolution 921,88] {a:4,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
924. 55 <=> r2_hidden(sK5(sK2,sK0),sK2) [avatar definition]
926. r2_hidden(sK5(sK2,sK0),sK2) <- (55) [avatar component clause 924] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
927. 55 | 1 | ~16 [avatar split clause 922,195,86,924]
959. r2_hidden(sK5(sK2,sK0),sK0) <- (20, 55) [resolution 926,222] {a:5,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
963. 56 <=> r2_hidden(sK5(sK2,sK0),sK0) [avatar definition]
965. r2_hidden(sK5(sK2,sK0),sK0) <- (56) [avatar component clause 963] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
966. 56 | ~20 | ~55 [avatar split clause 959,924,221,963]
1002. sK0 = sK2 | ~r2_hidden(sK5(sK2,sK0),sK2) <- (56) [resolution 965,68] {a:6,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
1007. ~r2_hidden(sK5(sK2,sK0),sK2) <- (~1, 56) [subsumption resolution 1002,88] {a:6,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
1008. $false <- (~1, 55, 56) [subsumption resolution 1007,926] {a:6,w:0,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
1009. 1 | ~55 | ~56 [avatar contradiction clause 1008]
1011. k1_xboole_0 != sK2 | sK0 != sK2 | k1_xboole_0 = sK0 [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
1012. $false [avatar sat refutation 93,98,103,108,193,206,234,253,255,491,505,506,534,721,777,788,927,966,1009,1011] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:770,thDist:-770}
% SZS output end Proof for zfmisc_1__t110_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 718
% Active clauses: 129
% Passive clauses: 398
% Final active clauses: 80
% Final passive clauses: 209
% Main loop iterations started: 167
% 
% >>> Simplifying Inferences
% Duplicate literals: 10
% Fw subsumption resolutions: 13
% Fw demodulations: 9
% Bw demodulations: 25
% 
% >>> Deletion Inferences
% Simple tautologies: 16
% Equational tautologies: 5
% Forward subsumptions: 211
% Backward subsumptions: 16
% 
% >>> Generating Inferences
% Binary resolution: 207
% Unit resulting resolution: 12
% Factoring: 3
% Forward superposition: 174
% Backward superposition: 154
% Self superposition: 13
% 
% >>> AVATAR
% Split clauses: 23
% Split components: 50
% Sat splitting refutations: 21
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 91
% SAT solver unit clauses: 8
% SAT solver binary clauses: 35
% 
% Memory used [KB]: 5756
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
% Time measurement results:
% SAT solver time: 0.002 s
% term sharing: 0.003 s
% backward demodulation index maintenance: 0.001 s
% resolution: 0.001 s ( own 0.0000 s ) 
% unit resulting resolution: 0.002 s
% superposition: 0.002 s ( own 0.0000 s ) 
% other: 0.010 s ( own 0.002 s ) 

