i: [1,42,0,4,1,1,0,0] t52_zfmisc_1
a: [3,68,0,4,1,42]
s: 68
i: [1,43,0,7,1,1,0,0] t52_zfmisc_1
a: [3,63,0,7,1,43]
s: 63
i: [1,44,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,73,0,2,1,44]
s: 73
i: [1,45,0,2,1,0,0,255] fc1_xboole_0
a: [3,78,0,2,1,45]
s: 78
i: [1,58,0,2,1,0,0,2] rc2_xboole_0
a: [3,88,0,2,1,58]
s: 88
i: [1,59,0,2,1,0,0,2] rc1_xboole_0
a: [3,93,0,2,1,59]
s: 93
i: [1,47,0,3,1,0,0,1] fc2_xboole_0
s: 47
i: [1,46,0,3,1,0,0,255] d2_xboole_0
a: [3,83,0,3,1,46]
s: 83
r: 78
i: [1,48,0,5,1,0,0,255] t2_boole
s: 48
i: [1,49,0,5,2,0,0,255] t6_boole
s: 49
i: [1,53,0,5,2,0,0,2] t7_boole
s: 53
i: [1,54,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 54
d: [2,96,0,5,1,1,22,48,83]
s: 96
d: [2,97,0,5,2,1,22,49,83]
s: 97
i: [1,56,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 56
i: [1,57,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 57
i: [1,50,0,7,2,0,0,1] l20_zfmisc_1
s: 50
i: [1,52,0,7,3,0,0,1] t8_boole
s: 52
d: [2,101,1,3,1,1,40,50,68]
a: [3,106,1,3,1,101]
s: 106
i: [1,55,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 55
i: [1,51,0,10,2,0,0,1] l24_zfmisc_1
s: 51
d: [2,109,2,2,1,1,40,106,53]
a: [3,118,2,2,1,109]
s: 118
d: [2,122,1,7,1,1,44,63,55]
a: [3,129,1,7,1,122]
s: 129
d: [2,136,2,0,0,2,52,106,129,51]
e: 136
d: [2,137,2,8,2,1,44,129,51]
d: [2,138,2,3,1,1,19,137]
d: [2,140,2,0,0,2,21,138,106]
e: 140
d: [2,124,1,7,1,1,44,63,55]
f: 140,124,101,42,43
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t52_zfmisc_1
% SZS output start Proof for zfmisc_1__t52_zfmisc_1
1. ! [X0] : ! [X1] : (k1_tarski(X0) = k3_xboole_0(k1_tarski(X0),X1) | r1_xboole_0(k1_tarski(X0),X1)) [input t52_zfmisc_1]
2. ~! [X0] : ! [X1] : (k1_tarski(X0) = k3_xboole_0(k1_tarski(X0),X1) | r1_xboole_0(k1_tarski(X0),X1)) [negated conjecture 1]
12. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
20. ! [X0] : ! [X1] : (~r2_hidden(X0,X1) => r1_xboole_0(k1_tarski(X0),X1)) [input l20_zfmisc_1]
21. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => k1_tarski(X0) = k3_xboole_0(X1,k1_tarski(X0))) [input l24_zfmisc_1]
22. ~! [X0,X1] : (k1_tarski(X0) = k3_xboole_0(k1_tarski(X0),X1) | r1_xboole_0(k1_tarski(X0),X1)) [flattening 2]
23. ! [X0,X1] : (~r2_hidden(X0,X1) => r1_xboole_0(k1_tarski(X0),X1)) [flattening 20]
24. ! [X0,X1] : (r2_hidden(X0,X1) => k1_tarski(X0) = k3_xboole_0(X1,k1_tarski(X0))) [flattening 21]
28. ? [X0,X1] : (k1_tarski(X0) != k3_xboole_0(k1_tarski(X0),X1) & ~r1_xboole_0(k1_tarski(X0),X1)) [ennf transformation 22]
30. ! [X0,X1] : (r1_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1)) [ennf transformation 23]
31. ! [X0,X1] : (k1_tarski(X0) = k3_xboole_0(X1,k1_tarski(X0)) | ~r2_hidden(X0,X1)) [ennf transformation 24]
36. ? [X0,X1] : (k1_tarski(X0) != k3_xboole_0(k1_tarski(X0),X1) & ~r1_xboole_0(k1_tarski(X0),X1)) => (k1_tarski(sK0) != k3_xboole_0(k1_tarski(sK0),sK1) & ~r1_xboole_0(k1_tarski(sK0),sK1)) [choice axiom]
37. k1_tarski(sK0) != k3_xboole_0(k1_tarski(sK0),sK1) & ~r1_xboole_0(k1_tarski(sK0),sK1) [skolemisation 28,36]
42. ~r1_xboole_0(k1_tarski(sK0),sK1) [cnf transformation 37] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. k1_tarski(sK0) != k3_xboole_0(k1_tarski(sK0),sK1) [cnf transformation 37] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. r1_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1) [cnf transformation 30] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k1_tarski(X0) = k3_xboole_0(X1,k1_tarski(X0)) | ~r2_hidden(X0,X1) [cnf transformation 31] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 12] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. 1 <=> k1_tarski(sK0) = k3_xboole_0(k1_tarski(sK0),sK1) [avatar definition]
63. k1_tarski(sK0) != k3_xboole_0(k1_tarski(sK0),sK1) <- (~1) [avatar component clause 61] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
64. ~1 [avatar split clause 43,61]
66. 2 <=> r1_xboole_0(k1_tarski(sK0),sK1) [avatar definition]
68. ~r1_xboole_0(k1_tarski(sK0),sK1) <- (~2) [avatar component clause 66] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
69. ~2 [avatar split clause 42,66]
101. r2_hidden(sK0,sK1) <- (~2) [resolution 50,68] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
104. 8 <=> r2_hidden(sK0,sK1) [avatar definition]
106. r2_hidden(sK0,sK1) <- (8) [avatar component clause 104] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
107. 8 | 2 [avatar split clause 101,66,104]
124. k1_tarski(sK0) != k3_xboole_0(sK1,k1_tarski(sK0)) <- (~1) [superposition 63,55] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
127. 11 <=> k1_tarski(sK0) = k3_xboole_0(sK1,k1_tarski(sK0)) [avatar definition]
129. k1_tarski(sK0) != k3_xboole_0(sK1,k1_tarski(sK0)) <- (~11) [avatar component clause 127] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
131. ~11 | 1 [avatar split clause 124,61,127]
137. k1_tarski(sK0) != k1_tarski(sK0) | ~r2_hidden(sK0,sK1) <- (~11) [superposition 129,51] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
138. ~r2_hidden(sK0,sK1) <- (~11) [trivial inequality removal 137] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
140. $false <- (8, ~11) [subsumption resolution 138,106] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
141. ~8 | 11 [avatar contradiction clause 140]
142. $false [avatar sat refutation 64,69,107,131,141] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:56,thDist:-56}
% SZS output end Proof for zfmisc_1__t52_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 18
% Generated clauses: 53
% Active clauses: 21
% Passive clauses: 35
% Final active clauses: 20
% Final passive clauses: 12
% Main loop iterations started: 23
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 1
% Fw demodulations: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% Forward superposition: 5
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 7
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

