i: [1,90,0,11,3,1,0,0] t60_group_6
a: [3,151,0,2,1,90]
s: 151
i: [1,85,0,2,1,1,0,0] t60_group_6
a: [3,169,0,2,1,85]
s: 169
a: [3,144,0,4,1,90]
s: 144
i: [1,86,0,4,1,1,0,0] t60_group_6
a: [3,164,0,4,1,86]
s: 164
a: [3,147,0,5,1,90]
s: 147
i: [1,87,0,6,1,1,0,0] t60_group_6
a: [3,159,0,6,1,87]
s: 159
i: [1,137,0,2,1,0,0,1] rc1_funct_1
a: [3,174,0,2,1,137]
s: 174
i: [1,136,0,2,1,0,0,1] rc1_funct_1
a: [3,179,0,2,1,136]
s: 179
i: [1,140,0,2,1,0,0,1] rc2_funct_1
a: [3,184,0,2,1,140]
s: 184
i: [1,139,0,2,1,0,0,1] rc2_funct_1
a: [3,189,0,2,1,139]
s: 189
i: [1,138,0,2,1,0,0,1] rc2_funct_1
a: [3,194,0,2,1,138]
s: 194
i: [1,101,0,3,1,0,0,255] rc2_funct_2
s: 101
i: [1,104,0,3,1,0,0,255] rc2_funct_2
s: 104
i: [1,108,0,3,1,0,0,255] reflexivity_r1_tarski
s: 108
i: [1,99,0,4,1,0,0,1] existence_m1_subset_1
s: 99
i: [1,102,0,4,1,0,0,255] rc2_funct_2
s: 102
i: [1,103,0,4,1,0,0,255] rc2_funct_2
s: 103
i: [1,105,0,4,1,0,0,255] rc2_funct_2
s: 105
i: [1,120,0,4,1,0,0,2] rc6_funct_1
s: 120
i: [1,123,0,4,1,0,0,2] rc6_funct_1
s: 123
i: [1,125,0,4,1,0,0,2] rc1_partfun1
s: 125
i: [1,128,0,4,1,0,0,2] rc1_partfun1
s: 128
i: [1,130,0,4,1,0,0,1] rc1_funct_2
s: 130
i: [1,133,0,4,1,0,0,1] rc1_funct_2
s: 133
i: [1,106,0,5,1,0,0,255] rc2_funct_2
s: 106
i: [1,107,0,5,1,0,0,255] rc2_funct_2
s: 107
i: [1,121,0,5,1,0,0,2] rc6_funct_1
s: 121
i: [1,122,0,5,1,0,0,2] rc6_funct_1
s: 122
i: [1,126,0,5,1,0,0,2] rc1_partfun1
s: 126
i: [1,127,0,5,1,0,0,2] rc1_partfun1
s: 127
i: [1,131,0,5,1,0,0,1] rc1_funct_2
s: 131
i: [1,132,0,5,1,0,0,1] rc1_funct_2
s: 132
i: [1,134,0,6,1,0,0,1] rc1_funct_2
s: 134
i: [1,95,0,7,2,0,0,255] t3_subset
s: 95
i: [1,100,0,7,1,0,0,255] rc2_funct_2
s: 100
i: [1,111,0,8,2,0,0,1] cc1_relset_1
s: 111
d: [2,197,1,2,1,1,40,111,159]
a: [3,203,1,2,1,197]
s: 203
i: [1,124,0,8,1,0,0,2] rc1_partfun1
s: 124
d: [2,198,1,6,1,0,40,111,99]
s: 198
i: [1,129,0,8,1,0,0,1] rc1_funct_2
s: 129
i: [1,112,0,9,2,0,0,2] cc2_relset_1
s: 112
d: [2,208,1,3,1,1,40,112,159]
a: [3,216,1,3,1,208]
s: 216
i: [1,113,0,9,2,0,0,2] cc2_relset_1
s: 113
d: [2,219,1,3,1,1,40,113,159]
a: [3,227,1,3,1,219]
s: 227
i: [1,91,0,10,4,0,0,1] cc3_funct_1
s: 91
i: [1,109,0,11,3,0,0,255] redefinition_k2_relset_1
s: 109
d: [2,235,1,9,3,1,44,109,147]
d: [2,237,1,6,2,2,21,235,227]
d: [2,238,1,4,1,3,21,237,203]
a: [3,242,1,4,1,238]
s: 242
i: [1,110,0,11,3,0,0,1] dt_k2_relset_1
s: 110
r: 144
r: 147
a: [3,143,0,4,1,90]
s: 143
a: [3,148,0,5,1,90]
s: 148
d: [2,261,1,9,3,1,44,148,109]
d: [2,262,1,5,2,2,21,261,242]
d: [2,263,1,2,1,3,21,262,227]
d: [2,264,1,0,0,4,21,263,203]
e: 264
r: 143
r: 148
s: 144
d: [2,251,1,9,3,1,44,110,147]
d: [2,254,1,6,2,2,21,251,227]
d: [2,255,1,4,1,3,21,254,203]
a: [3,259,1,4,1,255]
s: 259
s: 147
d: [2,196,1,7,2,0,40,111,95]
s: 196
i: [1,135,0,11,2,0,0,255] fc9_relset_1
s: 135
i: [1,92,0,13,3,0,0,255] cc4_funct_2
s: 92
d: [2,270,2,4,1,0,40,196,108]
s: 270
i: [1,93,0,13,4,0,0,1] d3_funct_2
s: 93
d: [2,282,1,11,2,0,40,92,129]
d: [2,283,1,5,1,0,21,282,134]
s: 283
i: [1,94,0,13,4,0,0,1] d3_funct_2
s: 94
i: [1,114,0,13,3,0,0,255] cc1_funct_2
s: 114
d: [2,292,1,11,4,1,44,94,147]
d: [2,293,1,8,3,1,19,292]
d: [2,296,1,5,2,2,21,293,227]
d: [2,297,1,3,1,3,21,296,203]
a: [3,301,1,3,1,297]
s: 301
i: [1,118,0,14,4,0,0,2] cc5_funct_2
s: 118
d: [2,314,1,11,3,0,40,118,100]
d: [2,317,1,6,2,0,21,314,107]
d: [2,318,1,3,1,0,21,317,104]
s: 318
i: [1,119,0,15,4,0,0,2] cc5_funct_2
s: 119
i: [1,96,0,16,6,0,0,2] d4_funct_2
s: 96
d: [2,325,1,12,3,0,40,119,100]
d: [2,328,1,7,2,0,21,325,107]
d: [2,329,1,4,1,0,21,328,104]
s: 329
i: [1,97,0,17,6,0,0,2] d4_funct_2
s: 97
d: [2,274,1,6,1,0,40,135,100]
s: 274
i: [1,116,0,17,5,0,0,2] cc6_funct_2
s: 116
d: [2,334,2,0,0,5,52,151,169,301,144,159,116]
e: 334
d: [2,337,1,11,4,1,40,116,159]
d: [2,343,1,8,3,2,21,337,301]
d: [2,344,1,6,2,3,21,343,151]
d: [2,345,1,4,1,4,21,344,169]
d: [2,346,1,0,0,5,21,345,144]
e: 346
r: 144
r: 147
r: 242
s: 143
s: 148
d: [2,353,2,0,0,4,52,203,227,301,148,93]
e: 353
d: [2,354,1,11,4,1,44,148,93]
d: [2,356,1,8,3,1,19,354]
d: [2,358,1,5,2,2,21,356,301]
d: [2,359,1,2,1,3,21,358,227]
d: [2,360,1,0,0,4,21,359,203]
e: 360
r: 148
r: 151
r: 301
i: [1,89,0,6,2,1,0,0] t60_group_6
i: [1,88,0,9,2,1,0,0] t60_group_6
d: [2,312,1,8,3,1,40,118,159]
d: [2,369,1,4,2,2,21,312,143]
d: [2,370,1,2,1,3,21,369,169]
d: [2,323,1,9,3,1,40,119,159]
d: [2,377,1,5,2,2,21,323,143]
d: [2,378,1,3,1,3,21,377,169]
f: 378,370,360,345,297,219,197,85,87,88,89,90
% Refutation found. Thanks to Tanya!
% SZS status Theorem for group_6__t60_group_6
% SZS output start Proof for group_6__t60_group_6
1. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X2,X0,X1) & v1_funct_1(X2)) => (v3_funct_2(X2,X0,X1) <=> (v2_funct_1(X2) & k2_relset_1(X1,X2) = X1))) [input t60_group_6]
2. ~! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X2,X0,X1) & v1_funct_1(X2)) => (v3_funct_2(X2,X0,X1) <=> (v2_funct_1(X2) & k2_relset_1(X1,X2) = X1))) [negated conjecture 1]
4. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [input cc1_relset_1]
5. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [input cc2_relset_1]
8. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ((v3_funct_2(X2,X0,X1) & v1_funct_1(X2)) => (v2_funct_2(X2,X1) & v2_funct_1(X2) & v1_funct_1(X2)))) [input cc5_funct_2]
9. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ((v2_funct_2(X2,X1) & v2_funct_1(X2) & v1_funct_1(X2)) => (v3_funct_2(X2,X0,X1) & v1_funct_1(X2)))) [input cc6_funct_2]
10. ! [X0] : ! [X1] : ((v5_relat_1(X1,X0) & v1_relat_1(X1)) => (v2_funct_2(X1,X0) <=> k2_relset_1(X0,X1) = X0)) [input d3_funct_2]
28. ~! [X0,X1,X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X2,X0,X1) & v1_funct_1(X2)) => (v3_funct_2(X2,X0,X1) <=> (v2_funct_1(X2) & k2_relset_1(X1,X2) = X1))) [flattening 2]
30. ! [X0,X1] : ((v5_relat_1(X1,X0) & v1_relat_1(X1)) => (v2_funct_2(X1,X0) <=> k2_relset_1(X0,X1) = X0)) [flattening 10]
34. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [flattening 4]
35. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [flattening 5]
37. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ((v2_funct_2(X2,X1) & v2_funct_1(X2) & v1_funct_1(X2)) => (v3_funct_2(X2,X0,X1) & v1_funct_1(X2)))) [flattening 9]
38. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => ((v3_funct_2(X2,X0,X1) & v1_funct_1(X2)) => (v2_funct_2(X2,X1) & v2_funct_1(X2) & v1_funct_1(X2)))) [flattening 8]
40. ? [X0,X1,X2] : ((v3_funct_2(X2,X0,X1) <~> (v2_funct_1(X2) & k2_relset_1(X1,X2) = X1)) & (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X2,X0,X1) & v1_funct_1(X2))) [ennf transformation 28]
41. ? [X0,X1,X2] : ((v3_funct_2(X2,X0,X1) <~> (v2_funct_1(X2) & k2_relset_1(X1,X2) = X1)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X2,X0,X1) & v1_funct_1(X2)) [flattening 40]
46. ! [X0,X1] : ((v2_funct_2(X1,X0) <=> k2_relset_1(X0,X1) = X0) | (~v5_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 30]
47. ! [X0,X1] : ((v2_funct_2(X1,X0) <=> k2_relset_1(X0,X1) = X0) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 46]
55. ! [X0,X1,X2] : (v1_relat_1(X2) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 34]
56. ! [X0,X1,X2] : ((v5_relat_1(X2,X1) & v4_relat_1(X2,X0)) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 35]
59. ! [X0,X1,X2] : (((v3_funct_2(X2,X0,X1) & v1_funct_1(X2)) | (~v2_funct_2(X2,X1) | ~v2_funct_1(X2) | ~v1_funct_1(X2))) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 37]
60. ! [X0,X1,X2] : ((v3_funct_2(X2,X0,X1) & v1_funct_1(X2)) | ~v2_funct_2(X2,X1) | ~v2_funct_1(X2) | ~v1_funct_1(X2) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [flattening 59]
61. ! [X0,X1,X2] : (((v2_funct_2(X2,X1) & v2_funct_1(X2) & v1_funct_1(X2)) | (~v3_funct_2(X2,X0,X1) | ~v1_funct_1(X2))) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 38]
62. ! [X0,X1,X2] : ((v2_funct_2(X2,X1) & v2_funct_1(X2) & v1_funct_1(X2)) | ~v3_funct_2(X2,X0,X1) | ~v1_funct_1(X2) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [flattening 61]
64. ? [X0,X1,X2] : ((((~v2_funct_1(X2) | k2_relset_1(X1,X2) != X1) | ~v3_funct_2(X2,X0,X1)) & ((v2_funct_1(X2) & k2_relset_1(X1,X2) = X1) | v3_funct_2(X2,X0,X1))) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X2,X0,X1) & v1_funct_1(X2)) [nnf transformation 41]
65. ? [X0,X1,X2] : ((~v2_funct_1(X2) | k2_relset_1(X1,X2) != X1 | ~v3_funct_2(X2,X0,X1)) & ((v2_funct_1(X2) & k2_relset_1(X1,X2) = X1) | v3_funct_2(X2,X0,X1)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X2,X0,X1) & v1_funct_1(X2)) [flattening 64]
66. ? [X0,X1,X2] : ((~v2_funct_1(X2) | k2_relset_1(X1,X2) != X1 | ~v3_funct_2(X2,X0,X1)) & ((v2_funct_1(X2) & k2_relset_1(X1,X2) = X1) | v3_funct_2(X2,X0,X1)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X2,X0,X1) & v1_funct_1(X2)) => ((~v2_funct_1(sK2) | sK1 != k2_relset_1(sK1,sK2) | ~v3_funct_2(sK2,sK0,sK1)) & ((v2_funct_1(sK2) & sK1 = k2_relset_1(sK1,sK2)) | v3_funct_2(sK2,sK0,sK1)) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(sK2,sK0,sK1) & v1_funct_1(sK2)) [choice axiom]
67. (~v2_funct_1(sK2) | sK1 != k2_relset_1(sK1,sK2) | ~v3_funct_2(sK2,sK0,sK1)) & ((v2_funct_1(sK2) & sK1 = k2_relset_1(sK1,sK2)) | v3_funct_2(sK2,sK0,sK1)) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_2(sK2,sK0,sK1) & v1_funct_1(sK2) [skolemisation 65,66]
68. ! [X0,X1] : (((v2_funct_2(X1,X0) | k2_relset_1(X0,X1) != X0) & (k2_relset_1(X0,X1) = X0 | ~v2_funct_2(X1,X0))) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 47]
85. v1_funct_1(sK2) [cnf transformation 67] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
87. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [cnf transformation 67] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
88. sK1 = k2_relset_1(sK1,sK2) | v3_funct_2(sK2,sK0,sK1) [cnf transformation 67] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
89. v2_funct_1(sK2) | v3_funct_2(sK2,sK0,sK1) [cnf transformation 67] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
90. ~v2_funct_1(sK2) | sK1 != k2_relset_1(sK1,sK2) | ~v3_funct_2(sK2,sK0,sK1) [cnf transformation 67] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
93. k2_relset_1(X0,X1) = X0 | ~v2_funct_2(X1,X0) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 68] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
94. k2_relset_1(X0,X1) != X0 | v2_funct_2(X1,X0) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 68] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
111. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v1_relat_1(X2) [cnf transformation 55] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
113. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v5_relat_1(X2,X1) [cnf transformation 56] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
116. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~v2_funct_2(X2,X1) | ~v2_funct_1(X2) | ~v1_funct_1(X2) | v3_funct_2(X2,X0,X1) [cnf transformation 60] {a:0,w:17,wCS:6800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
118. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~v3_funct_2(X2,X0,X1) | ~v1_funct_1(X2) | v2_funct_1(X2) [cnf transformation 62] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
119. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~v3_funct_2(X2,X0,X1) | ~v1_funct_1(X2) | v2_funct_2(X2,X1) [cnf transformation 62] {a:0,w:15,wCS:6000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
142. 1 <=> v3_funct_2(sK2,sK0,sK1) [avatar definition]
143. v3_funct_2(sK2,sK0,sK1) <- (1) [avatar component clause 142] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
146. 2 <=> sK1 = k2_relset_1(sK1,sK2) [avatar definition]
147. sK1 = k2_relset_1(sK1,sK2) <- (2) [avatar component clause 146] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
148. sK1 != k2_relset_1(sK1,sK2) <- (~2) [avatar component clause 146] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
150. 3 <=> v2_funct_1(sK2) [avatar definition]
151. v2_funct_1(sK2) <- (3) [avatar component clause 150] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
153. ~1 | ~2 | ~3 [avatar split clause 90,150,146,142]
154. 1 | 3 [avatar split clause 89,150,142]
155. 1 | 2 [avatar split clause 88,146,142]
157. 4 <=> m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [avatar definition]
159. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) <- (4) [avatar component clause 157] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
160. 4 [avatar split clause 87,157]
167. 6 <=> v1_funct_1(sK2) [avatar definition]
169. v1_funct_1(sK2) <- (6) [avatar component clause 167] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
170. 6 [avatar split clause 85,167]
197. v1_relat_1(sK2) <- (4) [resolution 111,159] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
201. 12 <=> v1_relat_1(sK2) [avatar definition]
203. v1_relat_1(sK2) <- (12) [avatar component clause 201] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
204. 12 | ~4 [avatar split clause 197,157,201]
219. v5_relat_1(sK2,sK1) <- (4) [resolution 113,159] {a:1,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
225. 14 <=> v5_relat_1(sK2,sK1) [avatar definition]
227. v5_relat_1(sK2,sK1) <- (14) [avatar component clause 225] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
228. 14 | ~4 [avatar split clause 219,157,225]
292. sK1 != sK1 | v2_funct_2(sK2,sK1) | ~v5_relat_1(sK2,sK1) | ~v1_relat_1(sK2) <- (2) [superposition 94,147] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
293. v2_funct_2(sK2,sK1) | ~v5_relat_1(sK2,sK1) | ~v1_relat_1(sK2) <- (2) [trivial inequality removal 292] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
296. v2_funct_2(sK2,sK1) | ~v1_relat_1(sK2) <- (2, 14) [subsumption resolution 293,227] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
297. v2_funct_2(sK2,sK1) <- (2, 12, 14) [subsumption resolution 296,203] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
299. 17 <=> v2_funct_2(sK2,sK1) [avatar definition]
301. v2_funct_2(sK2,sK1) <- (17) [avatar component clause 299] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
302. 17 | ~2 | ~12 | ~14 [avatar split clause 297,225,201,146,299]
312. ~v3_funct_2(sK2,sK0,sK1) | ~v1_funct_1(sK2) | v2_funct_1(sK2) <- (4) [resolution 118,159] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
323. ~v3_funct_2(sK2,sK0,sK1) | ~v1_funct_1(sK2) | v2_funct_2(sK2,sK1) <- (4) [resolution 119,159] {a:1,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
337. ~v2_funct_2(sK2,sK1) | ~v2_funct_1(sK2) | ~v1_funct_1(sK2) | v3_funct_2(sK2,sK0,sK1) <- (4) [resolution 116,159] {a:1,w:11,wCS:1100,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
343. ~v2_funct_1(sK2) | ~v1_funct_1(sK2) | v3_funct_2(sK2,sK0,sK1) <- (4, 17) [subsumption resolution 337,301] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
344. ~v1_funct_1(sK2) | v3_funct_2(sK2,sK0,sK1) <- (3, 4, 17) [subsumption resolution 343,151] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
345. v3_funct_2(sK2,sK0,sK1) <- (3, 4, 6, 17) [subsumption resolution 344,169] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
350. 1 | ~3 | ~4 | ~6 | ~17 [avatar split clause 345,299,167,157,150,142]
354. sK1 != sK1 | ~v2_funct_2(sK2,sK1) | ~v5_relat_1(sK2,sK1) | ~v1_relat_1(sK2) <- (~2) [superposition 148,93] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
356. ~v2_funct_2(sK2,sK1) | ~v5_relat_1(sK2,sK1) | ~v1_relat_1(sK2) <- (~2) [trivial inequality removal 354] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
358. ~v5_relat_1(sK2,sK1) | ~v1_relat_1(sK2) <- (~2, 17) [subsumption resolution 356,301] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
359. ~v1_relat_1(sK2) <- (~2, 14, 17) [subsumption resolution 358,227] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
360. $false <- (~2, 12, 14, 17) [subsumption resolution 359,203] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
361. 2 | ~12 | ~14 | ~17 [avatar contradiction clause 360]
369. ~v1_funct_1(sK2) | v2_funct_1(sK2) <- (1, 4) [subsumption resolution 312,143] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
370. v2_funct_1(sK2) <- (1, 4, 6) [subsumption resolution 369,169] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
371. 3 | ~1 | ~4 | ~6 [avatar split clause 370,167,157,142,150]
377. ~v1_funct_1(sK2) | v2_funct_2(sK2,sK1) <- (1, 4) [subsumption resolution 323,143] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
378. v2_funct_2(sK2,sK1) <- (1, 4, 6) [subsumption resolution 377,169] {a:1,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
379. 17 | ~1 | ~4 | ~6 [avatar split clause 378,167,157,142,299]
380. $false [avatar sat refutation 153,154,155,160,170,204,228,302,350,361,371,379] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:153,thDist:-153}
% SZS output end Proof for group_6__t60_group_6
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% >>> Preprocessing
% Introduced skolems: 10
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 56
% Generated clauses: 240
% Active clauses: 72
% Passive clauses: 120
% Final active clauses: 62
% Final passive clauses: 46
% Main loop iterations started: 72
% 
% >>> Simplifying Inferences
% Duplicate literals: 12
% Trivial inequalities: 2
% Fw subsumption resolutions: 43
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 31
% 
% >>> Generating Inferences
% Binary resolution: 74
% Unit resulting resolution: 2
% Forward superposition: 11
% Backward superposition: 5
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 30
% SAT solver unit clauses: 8
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9594
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% resolution: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

