i: [1,56,0,2,1,1,0,0] t43_xxreal_2
a: [3,75,0,2,1,56]
s: 75
i: [1,55,0,2,1,1,0,0] t43_xxreal_2
a: [3,80,0,2,1,55]
s: 80
i: [1,53,0,2,1,1,0,0] t43_xxreal_2
a: [3,90,0,2,1,53]
s: 90
i: [1,52,0,2,1,1,0,0] t43_xxreal_2
a: [3,95,0,2,1,52]
s: 95
i: [1,54,0,3,1,1,0,0] t43_xxreal_2
a: [3,85,0,3,1,54]
s: 85
i: [1,71,0,2,1,0,0,2] rc1_xreal_0
a: [3,100,0,2,1,71]
s: 100
i: [1,57,0,3,1,0,0,255] fc1_finsub_1
s: 57
i: [1,70,0,3,1,0,0,1] reflexivity_r1_tarski
s: 70
i: [1,58,0,4,2,0,0,255] cc1_finsub_1
s: 58
i: [1,59,0,4,2,0,0,255] cc1_finsub_1
s: 59
i: [1,69,0,4,1,0,0,4] existence_m1_subset_1
s: 69
d: [2,102,1,3,1,0,40,58,57]
s: 102
i: [1,65,0,6,2,0,0,2] existence_m1_xxreal_2
s: 65
d: [2,103,1,3,1,0,40,59,57]
s: 103
i: [1,66,0,6,3,0,0,255] cc2_finsub_1
s: 66
i: [1,60,0,7,3,0,0,1] d10_xxreal_2
s: 60
i: [1,67,0,7,2,0,0,255] t3_subset
s: 67
i: [1,68,0,7,2,0,0,255] t3_subset
s: 68
d: [2,105,1,5,1,0,40,67,69]
s: 105
i: [1,61,0,8,3,0,0,1] d10_xxreal_2
s: 61
i: [1,64,0,8,3,0,0,255] cc20_membered
s: 64
d: [2,107,1,6,2,0,40,64,69]
s: 107
i: [1,62,0,9,4,0,0,1] d10_xxreal_2
s: 62
d: [2,108,1,7,3,0,40,64,68]
s: 108
i: [1,63,0,13,5,0,0,1] t6_xxreal_2
s: 63
d: [2,116,0,11,4,0,21,63,108]
s: 116
d: [2,109,1,9,4,0,40,62,65]
d: [2,112,1,7,3,0,18,109]
s: 112
d: [2,117,1,11,4,0,40,116,65]
d: [2,120,1,9,3,0,18,117]
s: 120
d: [2,122,2,12,5,0,40,120,62]
d: [2,124,2,10,4,0,21,122,108]
s: 124
d: [2,118,1,13,5,0,40,116,61]
d: [2,119,1,11,4,0,18,118]
s: 119
d: [2,126,2,14,6,0,40,119,62]
d: [2,128,2,12,5,0,21,126,108]
d: [2,129,2,9,4,0,21,128,60]
s: 129
d: [2,130,3,0,0,4,52,75,80,90,85,129]
e: 130
d: [2,131,3,6,3,1,40,129,85]
d: [2,135,3,4,2,2,21,131,90]
d: [2,136,3,2,1,3,21,135,80]
d: [2,137,3,0,0,4,21,136,75]
e: 137
f: 137,53,54,55,56
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_2__t43_xxreal_2
% SZS output start Proof for xxreal_2__t43_xxreal_2
1. ! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => ((v4_xxreal_2(X1) & r1_tarski(X0,X1)) => v4_xxreal_2(X0)))) [input t43_xxreal_2]
2. ~! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => ((v4_xxreal_2(X1) & r1_tarski(X0,X1)) => v4_xxreal_2(X0)))) [negated conjecture 1]
5. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v2_membered(X1))) [input cc20_membered]
9. ! [X0] : (v2_membered(X0) => (v4_xxreal_2(X0) <=> ? [X1] : (m1_xxreal_2(X1,X0) & v1_xreal_0(X1)))) [input d10_xxreal_2]
19. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
20. ! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => (r1_tarski(X0,X1) => ! [X2] : (m1_xxreal_2(X2,X1) => m1_xxreal_2(X2,X0))))) [input t6_xxreal_2]
21. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 19]
28. ? [X0] : (? [X1] : ((~v4_xxreal_2(X0) & (v4_xxreal_2(X1) & r1_tarski(X0,X1))) & v2_membered(X1)) & v2_membered(X0)) [ennf transformation 2]
29. ? [X0] : (? [X1] : (~v4_xxreal_2(X0) & v4_xxreal_2(X1) & r1_tarski(X0,X1) & v2_membered(X1)) & v2_membered(X0)) [flattening 28]
31. ! [X0] : ((v4_xxreal_2(X0) <=> ? [X1] : (m1_xxreal_2(X1,X0) & v1_xreal_0(X1))) | ~v2_membered(X0)) [ennf transformation 9]
32. ! [X0] : (! [X1] : ((! [X2] : (m1_xxreal_2(X2,X0) | ~m1_xxreal_2(X2,X1)) | ~r1_tarski(X0,X1)) | ~v2_membered(X1)) | ~v2_membered(X0)) [ennf transformation 20]
33. ! [X0] : (! [X1] : (! [X2] : (m1_xxreal_2(X2,X0) | ~m1_xxreal_2(X2,X1)) | ~r1_tarski(X0,X1) | ~v2_membered(X1)) | ~v2_membered(X0)) [flattening 32]
34. ! [X0] : (! [X1] : (v2_membered(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v2_membered(X0)) [ennf transformation 5]
38. ? [X0] : (? [X1] : (~v4_xxreal_2(X0) & v4_xxreal_2(X1) & r1_tarski(X0,X1) & v2_membered(X1)) & v2_membered(X0)) => (? [X1] : (~v4_xxreal_2(sK0) & v4_xxreal_2(X1) & r1_tarski(sK0,X1) & v2_membered(X1)) & v2_membered(sK0)) [choice axiom]
39. ? [X1] : (~v4_xxreal_2(sK0) & v4_xxreal_2(X1) & r1_tarski(sK0,X1) & v2_membered(X1)) => (~v4_xxreal_2(sK0) & v4_xxreal_2(sK1) & r1_tarski(sK0,sK1) & v2_membered(sK1)) [choice axiom]
40. (~v4_xxreal_2(sK0) & v4_xxreal_2(sK1) & r1_tarski(sK0,sK1) & v2_membered(sK1)) & v2_membered(sK0) [skolemisation 29,39,38]
41. ! [X0] : (((v4_xxreal_2(X0) | ! [X1] : (~m1_xxreal_2(X1,X0) | ~v1_xreal_0(X1))) & (? [X1] : (m1_xxreal_2(X1,X0) & v1_xreal_0(X1)) | ~v4_xxreal_2(X0))) | ~v2_membered(X0)) [nnf transformation 31]
42. ! [X0] : (((v4_xxreal_2(X0) | ! [X1] : (~m1_xxreal_2(X1,X0) | ~v1_xreal_0(X1))) & (? [X2] : (m1_xxreal_2(X2,X0) & v1_xreal_0(X2)) | ~v4_xxreal_2(X0))) | ~v2_membered(X0)) [rectify 41]
43. ! [X0] : (? [X2] : (m1_xxreal_2(X2,X0) & v1_xreal_0(X2)) => (m1_xxreal_2(sK2(X0),X0) & v1_xreal_0(sK2(X0)))) [choice axiom]
44. ! [X0] : (((v4_xxreal_2(X0) | ! [X1] : (~m1_xxreal_2(X1,X0) | ~v1_xreal_0(X1))) & ((m1_xxreal_2(sK2(X0),X0) & v1_xreal_0(sK2(X0))) | ~v4_xxreal_2(X0))) | ~v2_membered(X0)) [skolemisation 42,43]
47. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 21]
53. v2_membered(sK1) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. r1_tarski(sK0,sK1) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. v4_xxreal_2(sK1) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. ~v4_xxreal_2(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
60. v1_xreal_0(sK2(X0)) | ~v4_xxreal_2(X0) | ~v2_membered(X0) [cnf transformation 44] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
61. m1_xxreal_2(sK2(X0),X0) | ~v4_xxreal_2(X0) | ~v2_membered(X0) [cnf transformation 44] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
62. ~m1_xxreal_2(X1,X0) | v4_xxreal_2(X0) | ~v1_xreal_0(X1) | ~v2_membered(X0) [cnf transformation 44] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
63. m1_xxreal_2(X2,X0) | ~m1_xxreal_2(X2,X1) | ~r1_tarski(X0,X1) | ~v2_membered(X1) | ~v2_membered(X0) [cnf transformation 33] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:1,thDist:-1}
64. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v2_membered(X1) | ~v2_membered(X0) [cnf transformation 34] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
68. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 47] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
73. 1 <=> v4_xxreal_2(sK0) [avatar definition]
75. ~v4_xxreal_2(sK0) <- (~1) [avatar component clause 73] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
76. ~1 [avatar split clause 56,73]
78. 2 <=> v4_xxreal_2(sK1) [avatar definition]
80. v4_xxreal_2(sK1) <- (2) [avatar component clause 78] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
81. 2 [avatar split clause 55,78]
83. 3 <=> r1_tarski(sK0,sK1) [avatar definition]
85. r1_tarski(sK0,sK1) <- (3) [avatar component clause 83] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. 3 [avatar split clause 54,83]
88. 4 <=> v2_membered(sK1) [avatar definition]
90. v2_membered(sK1) <- (4) [avatar component clause 88] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. 4 [avatar split clause 53,88]
108. ~r1_tarski(X1,X2) | ~v2_membered(X2) | v2_membered(X1) [resolution 64,68] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. ~m1_xxreal_2(X2,X1) | m1_xxreal_2(X2,X0) | ~r1_tarski(X0,X1) | ~v2_membered(X1) [subsumption resolution 63,108] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
118. m1_xxreal_2(sK2(X2),X3) | ~r1_tarski(X3,X2) | ~v2_membered(X2) | ~v4_xxreal_2(X2) | ~v2_membered(X2) [resolution 116,61] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:3,thDist:-3}
119. m1_xxreal_2(sK2(X2),X3) | ~r1_tarski(X3,X2) | ~v2_membered(X2) | ~v4_xxreal_2(X2) [duplicate literal removal 118] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
126. ~r1_tarski(X3,X4) | ~v2_membered(X4) | ~v4_xxreal_2(X4) | v4_xxreal_2(X3) | ~v1_xreal_0(sK2(X4)) | ~v2_membered(X3) [resolution 119,62] {a:2,w:14,wCS:5600,sine:1,thAx:0,allAx:5,thDist:-5}
128. ~r1_tarski(X3,X4) | ~v2_membered(X4) | ~v4_xxreal_2(X4) | v4_xxreal_2(X3) | ~v1_xreal_0(sK2(X4)) [subsumption resolution 126,108] {a:2,w:12,wCS:4800,sine:1,thAx:0,allAx:5,thDist:-5}
129. ~r1_tarski(X3,X4) | ~v2_membered(X4) | ~v4_xxreal_2(X4) | v4_xxreal_2(X3) [subsumption resolution 128,60] {a:2,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
131. ~v2_membered(sK1) | ~v4_xxreal_2(sK1) | v4_xxreal_2(sK0) <- (3) [resolution 129,85] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
135. ~v4_xxreal_2(sK1) | v4_xxreal_2(sK0) <- (3, 4) [subsumption resolution 131,90] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
136. v4_xxreal_2(sK0) <- (2, 3, 4) [subsumption resolution 135,80] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
137. $false <- (~1, 2, 3, 4) [subsumption resolution 136,75] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
138. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 137]
139. $false [avatar sat refutation 76,81,86,91,138] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:48,thDist:-48}
% SZS output end Proof for xxreal_2__t43_xxreal_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 19
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 20
% Generated clauses: 61
% Active clauses: 30
% Passive clauses: 33
% Final active clauses: 30
% Final passive clauses: 2
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 21
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 6
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

