i: [1,76,0,2,1,1,0,0] t4_isomichi
a: [3,118,0,2,1,76]
s: 118
i: [1,75,0,2,1,1,0,0] t4_isomichi
a: [3,123,0,2,1,75]
s: 123
i: [1,79,0,3,1,1,0,0] t4_isomichi
a: [3,103,0,3,1,79]
s: 103
i: [1,78,0,3,1,1,0,0] t4_isomichi
a: [3,108,0,3,1,78]
s: 108
i: [1,77,0,5,1,1,0,0] t4_isomichi
a: [3,113,0,5,1,77]
s: 113
i: [1,99,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,142,0,2,1,99]
s: 142
i: [1,93,0,3,1,0,0,255] reflexivity_r1_tarski
s: 93
i: [1,92,0,4,1,0,0,1] existence_m1_subset_1
s: 92
i: [1,91,0,7,2,0,0,255] t3_subset
s: 91
i: [1,86,0,8,3,0,0,1] rc10_pre_topc
s: 86
i: [1,88,0,8,3,0,0,1] rc1_tops_1
s: 88
i: [1,90,0,8,3,0,0,3] rc2_tops_1
s: 90
i: [1,82,0,24,7,0,0,2] t23_tops_1
a: [3,134,0,9,3,82]
s: 134
d: [2,144,1,0,0,4,52,118,123,113,134]
e: 144
d: [2,145,1,0,0,3,52,118,123,92,134]
e: 145
d: [2,146,1,4,2,2,40,134,113]
d: [2,151,1,2,1,3,21,146,123]
d: [2,152,1,0,0,4,21,151,118]
e: 152
r: 134
i: [1,84,0,10,3,0,0,3] rc6_pre_topc
s: 84
i: [1,85,0,10,3,0,0,1] rc10_pre_topc
s: 85
i: [1,87,0,10,3,0,0,1] rc1_tops_1
s: 87
i: [1,89,0,10,3,0,0,3] rc2_tops_1
s: 89
i: [1,94,0,14,4,0,0,2] fc9_tops_1
s: 94
d: [2,155,1,9,3,1,40,94,113]
d: [2,165,1,7,2,2,21,155,118]
d: [2,166,1,5,1,3,21,165,123]
a: [3,170,1,5,1,166]
s: 170
i: [1,95,0,14,3,0,0,2] dt_k2_pre_topc
s: 95
i: [1,96,0,14,3,0,0,2] dt_k1_tops_1
s: 96
d: [2,159,1,14,5,0,40,94,89]
d: [2,161,1,10,3,0,18,159]
s: 161
a: [3,137,0,15,4,82]
s: 137
d: [2,179,1,10,3,2,40,137,113]
d: [2,191,1,8,2,3,21,179,118]
d: [2,192,1,5,1,4,21,191,108]
a: [3,196,1,5,1,192]
s: 196
r: 170
i: [1,97,0,16,3,0,0,2] projectivity_k2_pre_topc
s: 97
i: [1,98,0,16,3,0,0,2] projectivity_k1_tops_1
s: 98
d: [2,158,1,14,5,0,40,94,87]
d: [2,162,1,10,3,0,18,158]
s: 162
i: [1,83,0,24,7,0,0,2] t23_tops_1
a: [3,130,0,17,5,83]
s: 130
d: [2,157,1,14,5,0,40,94,85]
d: [2,163,1,10,3,0,18,157]
s: 163
i: [1,80,0,19,4,0,0,1] d2_isomichi
s: 80
i: [1,81,0,19,4,0,0,1] d2_isomichi
s: 81
d: [2,221,2,17,4,1,44,81,196]
d: [2,227,2,10,3,1,19,221]
d: [2,228,2,7,2,2,21,227,103]
d: [2,229,2,2,1,3,21,228,113]
d: [2,230,2,0,0,4,21,229,118]
e: 230
f: 230,192,145,82,75,76,77,78,79
% Refutation found. Thanks to Tanya!
% SZS status Theorem for isomichi__t4_isomichi
% SZS output start Proof for isomichi__t4_isomichi
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v3_pre_topc(X1,X0) => v2_isomichi(X1,X0)))) [input t4_isomichi]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v3_pre_topc(X1,X0) => v2_isomichi(X1,X0)))) [negated conjecture 1]
4. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v2_isomichi(X1,X0) <=> k2_pre_topc(X0,k1_tops_1(X0,X1)) = k2_pre_topc(X0,X1)))) [input d2_isomichi]
14. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
24. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (l1_pre_topc(X1) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) => ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))) => ((k1_tops_1(X0,X2) = X2 => v3_pre_topc(X2,X0)) & (v3_pre_topc(X3,X1) => k1_tops_1(X1,X3) = X3)))))) [input t23_tops_1]
35. ? [X0] : (? [X1] : ((~v2_isomichi(X1,X0) & v3_pre_topc(X1,X0)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & (l1_pre_topc(X0) & v2_pre_topc(X0))) [ennf transformation 2]
36. ? [X0] : (? [X1] : (~v2_isomichi(X1,X0) & v3_pre_topc(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 35]
37. ! [X0] : (! [X1] : ((v2_isomichi(X1,X0) <=> k2_pre_topc(X0,k1_tops_1(X0,X1)) = k2_pre_topc(X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 4]
38. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((v3_pre_topc(X2,X0) | k1_tops_1(X0,X2) != X2) & (k1_tops_1(X1,X3) = X3 | ~v3_pre_topc(X3,X1))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))) | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X1)) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 24]
39. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((v3_pre_topc(X2,X0) | k1_tops_1(X0,X2) != X2) & (k1_tops_1(X1,X3) = X3 | ~v3_pre_topc(X3,X1))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))) | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X1)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 38]
59. ? [X0] : (? [X1] : (~v2_isomichi(X1,X0) & v3_pre_topc(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) => (? [X1] : (~v2_isomichi(X1,sK0) & v3_pre_topc(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0)) [choice axiom]
60. ? [X1] : (~v2_isomichi(X1,sK0) & v3_pre_topc(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => (~v2_isomichi(sK1,sK0) & v3_pre_topc(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
61. (~v2_isomichi(sK1,sK0) & v3_pre_topc(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) [skolemisation 36,60,59]
62. ! [X0] : (! [X1] : (((v2_isomichi(X1,X0) | k2_pre_topc(X0,k1_tops_1(X0,X1)) != k2_pre_topc(X0,X1)) & (k2_pre_topc(X0,k1_tops_1(X0,X1)) = k2_pre_topc(X0,X1) | ~v2_isomichi(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [nnf transformation 37]
71. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK6(X0),X0)) [choice axiom]
72. ! [X0] : m1_subset_1(sK6(X0),X0) [skolemisation 14,71]
75. v2_pre_topc(sK0) [cnf transformation 61] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
76. l1_pre_topc(sK0) [cnf transformation 61] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
77. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 61] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
78. v3_pre_topc(sK1,sK0) [cnf transformation 61] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
79. ~v2_isomichi(sK1,sK0) [cnf transformation 61] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
81. k2_pre_topc(X0,k1_tops_1(X0,X1)) != k2_pre_topc(X0,X1) | v2_isomichi(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 62] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
82. k1_tops_1(X1,X3) = X3 | ~v3_pre_topc(X3,X1) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))) | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X1) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 39] {a:0,w:24,wCS:9600,sine:2,thAx:0,allAx:1,thDist:-1}
92. m1_subset_1(sK6(X0),X0) [cnf transformation 72] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
101. 1 <=> v2_isomichi(sK1,sK0) [avatar definition]
103. ~v2_isomichi(sK1,sK0) <- (~1) [avatar component clause 101] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. ~1 [avatar split clause 79,101]
106. 2 <=> v3_pre_topc(sK1,sK0) [avatar definition]
108. v3_pre_topc(sK1,sK0) <- (2) [avatar component clause 106] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
109. 2 [avatar split clause 78,106]
111. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
113. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (3) [avatar component clause 111] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 3 [avatar split clause 77,111]
116. 4 <=> l1_pre_topc(sK0) [avatar definition]
118. l1_pre_topc(sK0) <- (4) [avatar component clause 116] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. 4 [avatar split clause 76,116]
121. 5 <=> v2_pre_topc(sK0) [avatar definition]
123. v2_pre_topc(sK0) <- (5) [avatar component clause 121] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. 5 [avatar split clause 75,121]
133. 8 <=> ! [X0,X2] : (~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) | ~v2_pre_topc(X0) | ~l1_pre_topc(X0)) [avatar definition]
134. ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) | ~v2_pre_topc(X0) | ~l1_pre_topc(X0) <- (8) [avatar component clause 133] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
136. 9 <=> ! [X1,X3] : (k1_tops_1(X1,X3) = X3 | ~l1_pre_topc(X1) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))) | ~v3_pre_topc(X3,X1)) [avatar definition]
137. ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))) | ~l1_pre_topc(X1) | k1_tops_1(X1,X3) = X3 | ~v3_pre_topc(X3,X1) <- (9) [avatar component clause 136] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
138. 8 | 9 [avatar split clause 82,136,133]
145. $false <- (4, 5, 8) [unit resulting resolution 118,123,92,134] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
150. ~4 | ~5 | ~8 [avatar contradiction clause 145]
179. ~l1_pre_topc(sK0) | sK1 = k1_tops_1(sK0,sK1) | ~v3_pre_topc(sK1,sK0) <- (3, 9) [resolution 137,113] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
191. sK1 = k1_tops_1(sK0,sK1) | ~v3_pre_topc(sK1,sK0) <- (3, 4, 9) [subsumption resolution 179,118] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
192. sK1 = k1_tops_1(sK0,sK1) <- (2, 3, 4, 9) [subsumption resolution 191,108] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
194. 12 <=> sK1 = k1_tops_1(sK0,sK1) [avatar definition]
196. sK1 = k1_tops_1(sK0,sK1) <- (12) [avatar component clause 194] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
197. 12 | ~2 | ~3 | ~4 | ~9 [avatar split clause 192,136,116,111,106,194]
221. k2_pre_topc(sK0,sK1) != k2_pre_topc(sK0,sK1) | v2_isomichi(sK1,sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (12) [superposition 81,196] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
227. v2_isomichi(sK1,sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (12) [trivial inequality removal 221] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
228. ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (~1, 12) [subsumption resolution 227,103] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
229. ~l1_pre_topc(sK0) <- (~1, 3, 12) [subsumption resolution 228,113] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
230. $false <- (~1, 3, 4, 12) [subsumption resolution 229,118] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
231. 1 | ~3 | ~4 | ~12 [avatar contradiction clause 230]
233. $false [avatar sat refutation 104,109,114,119,124,138,150,197,231] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:74,thDist:-74}
% SZS output end Proof for isomichi__t4_isomichi
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 3
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 25
% Generated clauses: 111
% Active clauses: 31
% Passive clauses: 44
% Final active clauses: 29
% Final passive clauses: 13
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Duplicate literals: 26
% Trivial inequalities: 4
% Fw subsumption resolutions: 13
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 21
% Unit resulting resolution: 2
% Forward superposition: 6
% Backward superposition: 5
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 6
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% splitting model update: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

