i: [1,184,0,2,1,1,0,0] t37_ordinal3
a: [3,293,0,2,1,184]
s: 293
i: [1,183,0,2,1,1,0,0] t37_ordinal3
a: [3,298,0,2,1,183]
s: 298
i: [1,186,0,6,2,1,0,0] t37_ordinal3
a: [3,279,0,3,1,186]
s: 279
i: [1,185,0,5,1,1,0,0] t37_ordinal3
a: [3,288,0,5,1,185]
s: 288
i: [1,187,0,2,1,0,0,1] spc1_boole
a: [3,303,0,2,1,187]
s: 303
i: [1,188,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,308,0,2,1,188]
s: 308
i: [1,189,0,2,1,0,0,255] spc0_boole
a: [3,313,0,2,1,189]
s: 313
i: [1,190,0,2,1,0,0,2] fc1_xboole_0
a: [3,318,0,2,1,190]
s: 318
i: [1,198,0,2,1,0,0,255] fc6_ordinal1
a: [3,353,0,2,1,198]
s: 353
i: [1,197,0,2,1,0,0,255] fc6_ordinal1
a: [3,358,0,2,1,197]
s: 358
i: [1,260,0,7,3,0,0,255] reflexivity_r1_ordinal1
a: [3,392,0,2,1,260]
s: 392
d: [2,450,1,0,0,2,40,392,353]
e: 450
d: [2,449,1,0,0,2,40,392,298]
e: 449
d: [2,448,1,0,0,2,40,392,293]
e: 448
d: [2,447,1,0,0,2,52,293,392]
e: 447
r: 392
i: [1,266,0,2,1,0,0,2] rc2_xboole_0
a: [3,400,0,2,1,266]
s: 400
i: [1,267,0,2,1,0,0,255] rc4_ordinal1
a: [3,405,0,2,1,267]
s: 405
i: [1,268,0,2,1,0,0,2] rc1_xboole_0
a: [3,410,0,2,1,268]
s: 410
i: [1,269,0,2,1,0,0,1] rc1_ordinal1
a: [3,415,0,2,1,269]
s: 415
i: [1,273,0,2,1,0,0,255] rc2_ordinal1
a: [3,420,0,2,1,273]
s: 420
i: [1,272,0,2,1,0,0,255] rc2_ordinal1
a: [3,425,0,2,1,272]
s: 425
i: [1,271,0,2,1,0,0,255] rc2_ordinal1
a: [3,430,0,2,1,271]
s: 430
i: [1,270,0,2,1,0,0,255] rc2_ordinal1
a: [3,435,0,2,1,270]
s: 435
i: [1,275,0,2,1,0,0,255] rc5_ordinal1
a: [3,440,0,2,1,275]
s: 440
i: [1,274,0,2,1,0,0,255] rc5_ordinal1
a: [3,445,0,2,1,274]
s: 445
i: [1,205,0,3,1,0,0,255] fc2_xboole_0
s: 205
i: [1,206,0,3,1,0,0,2] fc1_ordinal1
s: 206
i: [1,207,0,3,1,0,0,255] t2_xboole_1
s: 207
i: [1,248,0,3,1,0,0,255] irreflexivity_r2_xboole_0
s: 248
i: [1,249,0,3,1,0,0,255] reflexivity_r1_tarski
s: 249
i: [1,191,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,323,0,3,1,191]
s: 323
i: [1,192,0,3,1,0,0,255] d2_xboole_0
a: [3,328,0,3,1,192]
s: 328
r: 308
i: [1,193,0,3,1,0,0,255] t1_numerals
a: [3,333,0,3,1,193]
s: 333
i: [1,201,0,3,1,0,0,255] spc1_numerals
a: [3,363,0,3,1,201]
s: 363
i: [1,200,0,3,1,0,0,255] spc1_numerals
a: [3,368,0,3,1,200]
s: 368
i: [1,204,0,3,1,0,0,255] spc0_numerals
a: [3,378,0,3,1,204]
s: 378
i: [1,203,0,3,1,0,0,255] spc0_numerals
a: [3,383,0,3,1,203]
s: 383
d: [2,456,0,3,1,2,22,368,323]
a: [3,460,0,3,1,456]
s: 460
d: [2,462,0,3,1,2,22,383,323]
a: [3,466,0,3,1,462]
s: 466
i: [1,212,0,4,2,0,0,255] cc6_ordinal1
s: 212
d: [2,468,1,2,1,1,40,212,405]
a: [3,473,1,2,1,468]
s: 473
i: [1,213,0,4,2,0,0,255] cc9_ordinal1
s: 213
d: [2,469,1,2,1,1,40,212,440]
a: [3,478,1,2,1,469]
s: 478
i: [1,214,0,4,2,0,0,255] cc7_ordinal1
s: 214
i: [1,215,0,4,2,0,0,1] cc3_ordinal1
s: 215
d: [2,480,1,2,1,1,40,213,318]
a: [3,486,1,2,1,480]
s: 486
i: [1,219,0,4,2,0,0,255] cc1_ordinal1
s: 219
d: [2,535,1,2,1,1,40,219,298]
a: [3,549,1,2,1,535]
s: 549
i: [1,220,0,4,2,0,0,255] cc1_ordinal1
s: 220
i: [1,247,0,4,1,0,0,3] existence_m1_subset_1
s: 247
d: [2,536,1,2,1,1,40,219,293]
a: [3,554,1,2,1,536]
s: 554
i: [1,194,0,4,1,0,0,255] rqSucc__k1_ordinal1__r0_r1
a: [3,338,0,4,1,194]
s: 338
d: [2,572,1,2,1,1,40,220,298]
a: [3,586,1,2,1,572]
s: 586
i: [1,195,0,4,1,0,0,1] l14_ordinal3
a: [3,343,0,4,1,195]
s: 343
i: [1,196,0,4,1,0,0,255] dt_k5_numbers
a: [3,348,0,4,1,196]
s: 348
d: [2,573,1,2,1,1,40,220,293]
a: [3,591,1,2,1,573]
s: 591
i: [1,199,0,4,1,0,0,255] spc1_numerals
a: [3,373,0,4,1,199]
s: 373
d: [2,481,1,2,1,1,40,213,313]
a: [3,491,1,2,1,481]
s: 491
i: [1,202,0,4,1,0,0,255] spc0_numerals
a: [3,388,0,4,1,202]
s: 388
d: [2,610,0,4,1,2,22,348,323]
a: [3,614,0,4,1,610]
s: 614
d: [2,482,1,2,1,1,40,213,410]
a: [3,496,1,2,1,482]
s: 496
d: [2,616,0,4,1,2,22,373,323]
a: [3,620,0,4,1,616]
s: 620
d: [2,498,1,2,1,1,40,214,318]
a: [3,504,1,2,1,498]
s: 504
d: [2,622,0,4,1,2,22,388,323]
a: [3,626,0,4,1,622]
s: 626
i: [1,208,0,5,1,0,0,255] t1_boole
s: 208
d: [2,499,1,2,1,1,40,214,313]
a: [3,509,1,2,1,499]
s: 509
i: [1,216,0,5,2,0,0,2] t6_boole
s: 216
d: [2,500,1,2,1,1,40,214,410]
a: [3,514,1,2,1,500]
s: 514
i: [1,221,0,5,2,0,0,2] fc2_ordinal1
s: 221
i: [1,222,0,5,2,0,0,2] fc2_ordinal1
s: 222
d: [2,516,1,2,1,1,40,215,318]
a: [3,522,1,2,1,516]
s: 522
i: [1,234,0,5,2,0,0,255] cc8_ordinal1
s: 234
d: [2,517,1,2,1,1,40,215,313]
a: [3,527,1,2,1,517]
s: 527
i: [1,243,0,5,2,0,0,2] t7_boole
s: 243
i: [1,250,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 250
d: [2,518,1,2,1,1,40,215,410]
a: [3,532,1,2,1,518]
s: 532
a: [3,395,0,5,2,260]
s: 395
d: [2,534,1,2,1,1,40,219,353]
a: [3,544,1,2,1,534]
s: 544
i: [1,233,0,6,2,0,0,255] t3_xboole_1
s: 233
i: [1,235,0,6,3,0,0,255] cc2_ordinal1
s: 235
d: [2,538,1,2,1,1,40,219,415]
a: [3,564,1,2,1,538]
s: 564
i: [1,237,0,6,2,0,0,2] t1_subset
s: 237
d: [2,571,1,2,1,1,40,220,353]
a: [3,581,1,2,1,571]
s: 581
i: [1,244,0,6,2,0,0,255] t5_ordinal1
s: 244
i: [1,252,0,6,2,0,0,255] fc4_xboole_0
s: 252
d: [2,575,1,2,1,1,40,220,415]
a: [3,601,1,2,1,575]
s: 601
i: [1,253,0,6,2,0,0,255] fc5_xboole_0
s: 253
d: [2,644,1,4,2,1,44,222,338]
a: [3,649,1,2,1,644]
s: 649
i: [1,254,0,6,2,0,0,255] antisymmetry_r2_xboole_0
s: 254
i: [1,255,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 255
d: [2,665,1,2,1,1,40,234,460]
a: [3,672,1,2,1,665]
s: 672
i: [1,209,0,7,1,0,0,255] d1_ordinal1
s: 209
d: [2,537,2,2,1,1,40,219,473]
a: [3,559,2,2,1,537]
s: 559
i: [1,217,0,7,2,0,0,1] t38_ordinal2
s: 217
i: [1,218,0,7,2,0,0,1] t35_ordinal2
s: 218
d: [2,540,2,2,1,1,40,219,478]
a: [3,569,2,2,1,540]
s: 569
i: [1,223,0,7,2,0,0,1] t39_ordinal2
s: 223
d: [2,574,2,2,1,1,40,220,473]
a: [3,596,2,2,1,574]
s: 596
i: [1,224,0,7,2,0,0,1] t39_ordinal2
s: 224
i: [1,231,0,7,3,0,0,3] cc5_ordinal1
s: 231
d: [2,577,2,2,1,1,40,220,478]
a: [3,606,2,2,1,577]
s: 606
i: [1,232,0,7,3,0,0,3] cc2_ordinal2
s: 232
d: [2,652,2,2,1,1,40,522,220]
a: [3,658,2,2,1,652]
s: 658
i: [1,236,0,7,3,0,0,255] fc7_ordinal1
s: 236
d: [2,745,0,5,2,0,21,236,212]
s: 745
d: [2,653,2,2,1,1,40,522,219]
a: [3,663,2,2,1,653]
s: 663
i: [1,240,0,7,2,0,0,255] t3_subset
s: 240
d: [2,679,2,2,1,1,40,527,220]
a: [3,684,2,2,1,679]
s: 684
i: [1,241,0,7,2,0,0,255] t3_subset
s: 241
i: [1,242,0,7,3,0,0,1] t8_boole
s: 242
d: [2,680,2,2,1,1,40,527,219]
a: [3,689,2,2,1,680]
s: 689
i: [1,251,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 251
d: [2,691,2,2,1,1,40,532,220]
a: [3,696,2,2,1,691]
s: 696
i: [1,210,0,8,3,0,0,255] cc10_ordinal1
s: 210
i: [1,238,0,8,3,0,0,2] t2_subset
s: 238
d: [2,770,1,4,2,1,40,210,614]
a: [3,775,1,2,1,770]
s: 775
i: [1,261,0,8,3,0,0,1] dt_k11_ordinal2
s: 261
d: [2,782,1,5,2,1,40,238,363]
a: [3,796,1,2,1,782]
s: 796
d: [2,841,2,0,0,2,52,775,796,213]
e: 841
d: [2,844,2,2,1,1,40,796,213]
d: [2,846,2,0,0,2,21,844,775]
e: 846
r: 775
i: [1,262,0,8,3,0,0,255] fc1_ordinal3
s: 262
i: [1,239,0,9,3,0,0,255] d8_xboole_0
s: 239
a: [3,774,1,2,1,770]
s: 774
i: [1,246,0,9,3,0,0,255] t5_subset
s: 246
a: [3,779,1,2,1,770]
s: 779
i: [1,211,0,10,4,0,0,255] t11_ordinal1
s: 211
i: [1,245,0,10,3,0,0,255] t4_subset
s: 245
d: [2,692,2,2,1,1,40,532,219]
a: [3,701,2,2,1,692]
s: 701
i: [1,263,0,10,4,0,0,255] connectedness_r1_ordinal1
s: 263
d: [2,717,2,2,1,1,40,649,220]
a: [3,722,2,2,1,717]
s: 722
i: [1,264,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 264
i: [1,265,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 265
d: [2,718,2,2,1,1,40,649,219]
a: [3,727,2,2,1,718]
s: 727
i: [1,225,0,11,4,0,0,255] t21_ordinal1
s: 225
d: [2,843,2,2,1,1,40,796,214]
a: [3,852,2,2,1,843]
s: 852
i: [1,226,0,11,4,0,0,255] t21_ordinal1
s: 226
i: [1,227,0,11,4,0,0,255] t22_ordinal1
s: 227
d: [2,842,2,2,1,1,40,796,215]
a: [3,857,2,2,1,842]
s: 857
r: 796
r: 852
r: 857
i: [1,228,0,11,4,0,0,255] t22_ordinal1
s: 228
a: [3,795,1,2,1,782]
s: 795
i: [1,229,0,13,5,0,0,2] t14_ordinal1
s: 229
i: [1,259,0,14,4,0,0,255] existence_m2_subset_1
s: 259
a: [3,851,2,2,1,843]
s: 851
i: [1,256,0,15,5,0,0,255] dt_m2_subset_1
s: 256
a: [3,856,2,2,1,842]
s: 856
i: [1,257,0,15,5,0,0,255] redefinition_m2_subset_1
s: 257
i: [1,258,0,15,5,0,0,255] redefinition_m2_subset_1
s: 258
d: [2,668,1,3,1,0,40,234,247]
a: [3,677,1,3,1,668]
s: 677
i: [1,230,0,19,6,0,0,1] t40_ordinal2
s: 230
d: [2,972,1,17,6,1,44,230,288]
d: [2,995,1,15,5,2,21,972,298]
d: [2,996,1,13,4,3,21,995,293]
a: [3,1000,1,3,1,996]
s: 1000
r: 591
r: 554
r: 288
r: 293
r: 1000
s: 293
s: 288
s: 554
s: 591
a: [3,999,1,3,1,996]
s: 999
a: [3,1003,1,10,3,996]
s: 1003
d: [2,977,1,17,6,1,44,230,288]
d: [2,1012,1,15,5,2,21,977,298]
d: [2,1013,1,13,4,3,21,1012,293]
a: [3,1016,1,10,3,1013]
s: 1016
d: [2,1045,2,10,4,1,44,1003,223]
d: [2,1048,2,8,3,2,21,1045,649]
d: [2,1049,2,6,2,3,21,1048,293]
a: [3,1053,2,3,1,1049]
s: 1053
r: 279
a: [3,278,0,3,1,186]
s: 278
r: 288
r: 298
r: 549
r: 586
r: 1003
r: 1016
r: 1053
r: 278
r: 999
s: 298
s: 279
a: [3,283,0,3,1,186]
s: 283
s: 549
s: 288
s: 586
d: [2,1021,0,5,1,2,23,288,1000]
a: [3,1026,0,5,1,1021]
s: 1026
r: 1026
a: [3,1025,0,5,1,1021]
s: 1025
s: 999
d: [2,1157,1,7,2,1,44,1025,216]
s: 1157
d: [2,1161,2,5,2,2,44,1157,288]
d: [2,1164,2,2,1,2,19,1161]
a: [3,1168,2,2,1,1164]
s: 1168
s: 1003
s: 1016
a: [3,1052,2,3,1,1049]
s: 1052
r: 1052
d: [2,710,1,3,1,0,40,244,207]
s: 710
s: 1053
d: [2,711,1,3,1,0,40,244,249]
s: 711
d: [2,750,1,3,1,1,40,240,614]
a: [3,755,1,3,1,750]
s: 755
a: [3,1057,2,3,1,1049]
s: 1057
r: 1057
a: [3,792,1,3,1,782]
s: 792
a: [3,1056,2,3,1,1049]
s: 1056
d: [2,783,1,5,2,1,40,238,460]
d: [2,798,1,3,1,2,21,783,358]
a: [3,802,1,3,1,798]
s: 802
d: [2,785,1,5,2,1,40,238,333]
d: [2,813,1,3,1,2,21,785,358]
a: [3,817,1,3,1,813]
s: 817
d: [2,1046,2,10,4,1,44,1003,218]
d: [2,1059,2,8,3,2,21,1046,522]
d: [2,1060,2,6,2,3,21,1059,293]
a: [3,1064,2,3,1,1060]
s: 1064
d: [2,786,1,5,2,1,40,238,378]
a: [3,822,1,3,1,786]
s: 822
d: [2,1047,2,8,3,2,44,1003,288]
d: [2,1070,2,6,2,3,21,1047,298]
d: [2,1071,2,3,1,3,21,1070,255]
a: [3,1075,2,3,1,1071]
s: 1075
d: [2,787,1,5,2,1,40,238,466]
d: [2,824,1,3,1,2,21,787,358]
a: [3,828,1,3,1,824]
s: 828
d: [2,969,1,19,7,0,44,230,224]
d: [2,983,1,17,6,0,18,969]
d: [2,984,1,15,5,1,21,983,649]
a: [3,987,1,3,1,984]
s: 987
d: [2,1080,2,10,4,1,44,1016,223]
d: [2,1083,2,8,3,2,21,1080,649]
d: [2,1084,2,6,2,3,21,1083,293]
a: [3,1087,2,3,1,1084]
s: 1087
r: 1053
r: 1056
r: 1087
d: [2,784,1,7,2,1,40,238,614]
a: [3,807,1,4,1,784]
s: 807
s: 1052
a: [3,810,1,3,1,784]
s: 810
d: [2,633,1,5,2,0,44,207,216]
s: 633
s: 1057
d: [2,636,1,5,2,1,44,333,216]
s: 636
a: [3,1088,2,3,1,1084]
s: 1088
d: [2,642,1,5,2,0,40,222,220]
s: 642
d: [2,643,1,5,2,0,40,222,219]
s: 643
a: [3,1091,2,3,1,1084]
s: 1091
d: [2,741,1,5,2,0,40,231,247]
s: 741
d: [2,1081,2,10,4,1,44,1016,218]
d: [2,1094,2,8,3,2,21,1081,522]
d: [2,1095,2,6,2,3,21,1094,293]
a: [3,1099,2,3,1,1095]
s: 1099
r: 1099
d: [2,746,1,5,2,0,40,745,212]
s: 746
d: [2,751,1,5,1,0,40,240,247]
s: 751
a: [3,1098,2,3,1,1095]
s: 1098
r: 1098
d: [2,760,1,5,1,0,44,251,208]
s: 760
s: 1099
d: [2,761,1,5,1,0,44,251,208]
s: 761
d: [2,764,1,5,1,0,44,208,251]
s: 764
a: [3,1103,2,3,1,1095]
s: 1103
r: 279
d: [2,1116,0,5,1,2,23,288,278]
a: [3,1126,0,5,1,1116]
s: 1126
d: [2,1413,1,5,2,1,44,1126,223]
d: [2,1420,1,2,1,2,21,1413,283]
d: [2,1421,1,0,0,3,21,1420,293]
e: 1421
d: [2,1414,1,5,2,1,44,223,1126]
d: [2,1423,1,2,1,2,21,1414,283]
d: [2,1424,1,0,0,3,21,1423,293]
e: 1424
d: [2,1416,2,8,3,2,44,1003,1126]
d: [2,1426,2,5,2,2,21,1416,711]
d: [2,1427,2,3,1,3,21,1426,649]
d: [2,1428,2,0,0,4,21,1427,1052]
e: 1428
r: 1126
s: 279
d: [2,767,1,5,1,0,44,208,251]
s: 767
d: [2,883,1,7,3,0,40,265,207]
d: [2,886,1,5,2,1,21,883,522]
s: 886
d: [2,1204,3,3,1,1,40,1052,237]
a: [3,1210,3,3,1,1204]
s: 1210
d: [2,637,1,6,2,1,44,343,216]
s: 637
d: [2,1244,3,3,1,1,40,1057,237]
a: [3,1250,3,3,1,1244]
s: 1250
d: [2,771,1,6,2,0,40,210,247]
s: 771
d: [2,788,1,6,2,0,40,238,247]
s: 788
d: [2,1295,3,3,1,1,40,1087,237]
a: [3,1300,3,3,1,1295]
s: 1300
d: [2,867,1,6,2,0,40,239,207]
s: 867
d: [2,1296,3,2,1,1,40,1087,243]
a: [3,1304,3,2,1,1296]
s: 1304
d: [2,874,1,6,2,1,40,245,614]
s: 874
d: [2,897,1,12,4,1,40,226,395]
d: [2,902,1,9,3,1,18,897]
d: [2,904,1,6,2,1,21,902,222]
s: 904
d: [2,1473,4,2,1,1,40,1304,214]
a: [3,1478,4,2,1,1473]
s: 1478
d: [2,632,1,7,3,0,44,216,216]
s: 632
d: [2,1474,4,2,1,1,40,1304,213]
a: [3,1483,4,2,1,1474]
s: 1483
d: [2,634,1,7,2,0,44,208,216]
s: 634
d: [2,769,1,7,3,0,40,210,241]
s: 769
d: [2,1344,3,10,4,1,40,1099,229]
d: [2,1346,3,7,3,1,21,1344,710]
d: [2,1347,3,5,2,2,21,1346,522]
d: [2,1348,3,3,1,3,21,1347,298]
a: [3,1352,3,3,1,1348]
s: 1352
r: 288
r: 298
r: 549
r: 586
r: 1003
r: 1016
r: 1025
r: 1052
r: 1088
r: 1099
r: 1157
r: 1210
r: 1300
r: 1304
r: 1478
r: 1483
d: [2,1534,0,5,1,2,23,288,1352]
a: [3,1554,0,5,1,1534]
s: 1554
d: [2,1568,1,5,2,1,44,1554,218]
d: [2,1573,1,2,1,2,21,1568,987]
d: [2,1574,1,0,0,3,21,1573,293]
e: 1574
d: [2,1569,1,5,2,1,44,218,1554]
d: [2,1576,1,2,1,2,21,1569,987]
d: [2,1577,1,0,0,3,21,1576,293]
e: 1577
r: 1352
r: 1554
s: 298
s: 288
s: 1025
s: 549
a: [3,1553,0,5,1,1534]
s: 1553
s: 586
d: [2,1540,0,5,1,2,23,1025,1352]
a: [3,1559,0,5,1,1540]
s: 1559
s: 1157
a: [3,1305,3,2,1,1296]
s: 1305
d: [2,1597,1,7,2,1,44,1025,216]
s: 1597
s: 1478
d: [2,1599,1,7,2,1,44,1553,216]
s: 1599
d: [2,1602,1,7,2,1,44,1559,216]
s: 1602
s: 1483
s: 1003
s: 1052
s: 1016
d: [2,871,1,7,2,0,40,246,247]
s: 871
s: 1088
d: [2,955,1,11,4,1,40,258,614]
d: [2,957,1,9,3,2,21,955,358]
d: [2,958,1,7,2,3,21,957,795]
s: 958
s: 1098
d: [2,705,1,8,3,0,44,233,216]
s: 705
d: [2,759,1,8,3,0,40,241,231]
s: 759
s: 1210
d: [2,830,1,8,3,0,40,261,220]
s: 830
a: [3,1299,3,3,1,1295]
s: 1299
d: [2,831,1,8,3,0,40,261,219]
s: 831
d: [2,859,1,8,3,0,40,262,220]
s: 859
a: [3,1351,3,3,1,1348]
s: 1351
d: [2,860,1,8,3,0,40,262,219]
s: 860
d: [2,1368,3,3,1,1,40,1098,237]
a: [3,1374,3,3,1,1368]
s: 1374
d: [2,869,1,8,3,0,40,246,241]
s: 869
d: [2,875,1,8,2,0,40,245,247]
s: 875
d: [2,1402,3,3,1,1,40,1103,237]
a: [3,1408,3,3,1,1402]
s: 1408
d: [2,888,1,10,4,1,44,225,338]
d: [2,892,1,8,3,2,21,888,527]
s: 892
d: [2,1546,3,3,1,2,23,1300,1352]
a: [3,1565,3,3,1,1546]
s: 1565
d: [2,889,1,10,4,1,44,225,343]
d: [2,893,1,8,3,2,21,889,522]
s: 893
d: [2,899,1,10,4,1,44,226,338]
d: [2,906,1,8,3,2,21,899,527]
s: 906
d: [2,1738,3,5,2,3,40,893,1098]
d: [2,1746,3,3,1,4,21,1738,298]
a: [3,1750,3,3,1,1746]
s: 1750
d: [2,900,1,10,4,1,44,226,343]
d: [2,907,1,8,3,2,21,900,522]
s: 907
d: [2,1767,4,7,3,1,40,1750,264]
d: [2,1769,4,5,2,2,21,1767,298]
d: [2,1770,4,3,1,3,21,1769,649]
a: [3,1774,4,3,1,1770]
s: 1774
d: [2,1788,5,0,0,2,52,1052,1774,244]
e: 1788
d: [2,1790,5,3,1,1,40,1774,244]
d: [2,1798,5,0,0,2,21,1790,1052]
e: 1798
r: 283
r: 1052
r: 1057
r: 1088
r: 1091
r: 1299
r: 1210
r: 1250
r: 1483
d: [2,908,1,10,4,1,44,227,338]
d: [2,910,1,8,3,2,21,908,527]
s: 910
d: [2,909,1,10,4,1,44,227,343]
d: [2,911,1,8,3,2,21,909,522]
s: 911
a: [3,1482,4,2,1,1474]
s: 1482
d: [2,932,1,10,4,1,44,228,338]
d: [2,935,1,8,3,2,21,932,527]
s: 935
s: 1053
d: [2,933,1,10,4,1,44,228,343]
d: [2,936,1,8,3,2,21,933,522]
s: 936
d: [2,732,1,9,3,0,44,217,216]
s: 732
s: 1056
d: [2,733,1,9,3,0,44,218,216]
s: 733
s: 1087
d: [2,873,1,9,3,0,40,245,241]
s: 873
d: [2,781,1,10,3,0,40,238,241]
s: 781
a: [3,1092,2,3,1,1084]
s: 1092
d: [2,880,1,14,6,0,40,264,263]
d: [2,881,1,10,4,0,18,880]
s: 881
a: [3,1209,3,3,1,1204]
s: 1209
d: [2,939,1,12,5,0,40,229,243]
d: [2,944,1,10,4,0,21,939,215]
s: 944
d: [2,887,1,16,6,0,40,225,264]
d: [2,890,1,14,5,0,18,887]
d: [2,891,1,11,4,0,21,890,222]
s: 891
a: [3,1249,3,3,1,1244]
s: 1249
d: [2,898,1,16,6,0,40,226,263]
d: [2,901,1,14,5,0,18,898]
d: [2,905,1,11,4,0,21,901,222]
s: 905
s: 1300
d: [2,929,1,11,4,0,40,228,255]
s: 929
d: [2,930,1,11,4,0,40,228,237]
s: 930
d: [2,1827,3,5,2,3,40,910,1103]
d: [2,1831,3,3,1,4,21,1827,522]
a: [3,1835,3,3,1,1831]
s: 1835
a: [3,991,1,12,4,984]
s: 991
d: [2,1971,3,5,2,3,40,1092,911]
d: [2,1976,3,3,1,4,21,1971,293]
a: [3,1980,3,3,1,1976]
s: 1980
d: [2,974,1,19,7,0,44,230,224]
d: [2,980,1,17,6,0,18,974]
d: [2,1005,1,15,5,1,21,980,649]
a: [3,1008,1,12,4,1005]
s: 1008
d: [2,938,1,13,5,0,40,229,237]
s: 938
d: [2,1972,3,5,2,3,40,1092,910]
d: [2,1982,3,3,1,4,21,1972,293]
a: [3,1986,3,3,1,1982]
s: 1986
d: [2,948,1,21,7,0,40,256,259]
d: [2,949,1,13,4,0,18,948]
s: 949
d: [2,1974,3,3,1,1,40,1092,237]
a: [3,1991,3,3,1,1974]
s: 1991
d: [2,952,1,21,7,0,40,257,259]
d: [2,953,1,13,4,0,18,952]
s: 953
d: [2,954,1,14,5,0,40,258,241]
s: 954
d: [2,1768,4,5,2,3,40,1750,906]
d: [2,1776,4,3,1,4,21,1768,298]
a: [3,1780,4,3,1,1776]
s: 1780
d: [2,940,1,20,7,0,40,229,227]
d: [2,943,1,18,6,0,18,940]
d: [2,945,1,15,5,0,21,943,222]
s: 945
d: [2,2114,4,7,3,1,40,1980,264]
d: [2,2116,4,5,2,2,21,2114,522]
d: [2,2117,4,3,1,3,21,2116,293]
a: [3,2121,4,3,1,2117]
s: 2121
d: [2,2235,5,0,0,2,52,999,2121,233]
e: 2235
d: [2,2236,5,0,0,3,52,318,999,2121,705]
e: 2236
d: [2,2237,5,3,1,1,40,2121,233]
d: [2,2249,5,0,0,2,21,2237,999]
e: 2249
d: [2,2241,5,5,2,1,40,2121,705]
d: [2,2252,5,2,1,2,21,2241,999]
d: [2,2253,5,0,0,3,21,2252,318]
e: 2253
d: [2,645,1,4,2,1,44,222,343]
d: [2,1149,1,5,2,1,44,217,1026]
d: [2,1155,1,3,1,2,21,1149,298]
d: [2,1395,3,10,4,1,40,1099,229]
d: [2,1397,3,7,3,1,21,1395,710]
d: [2,1398,3,5,2,2,21,1397,522]
d: [2,1399,3,3,1,3,21,1398,298]
d: [2,1586,1,17,6,1,44,230,288]
d: [2,1589,1,14,5,2,21,1586,999]
d: [2,1590,1,12,4,3,21,1589,298]
d: [2,1591,1,10,3,4,21,1590,293]
d: [2,1634,2,10,4,1,44,1016,223]
d: [2,1811,2,8,3,2,21,1634,649]
d: [2,1812,2,6,2,3,21,1811,293]
d: [2,1870,3,10,4,1,40,1053,229]
d: [2,1871,3,7,3,2,21,1870,279]
d: [2,1872,3,5,2,3,21,1871,298]
d: [2,1873,3,3,1,4,21,1872,649]
f: 2249,2117,1976,1873,1812,1790,1770,1746,1591,1577,1534,1424,1399,1155,1116,1021,645,516,195,192,190,188,187,183,184,185,186
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t37_ordinal3
% SZS output start Proof for ordinal3__t37_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (k11_ordinal2(X0,X1) = 1 => (1 = X1 & 1 = X0)))) [input t37_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (k11_ordinal2(X0,X1) = 1 => (1 = X1 & 1 = X0)))) [negated conjecture 1]
9. ! [X0] : (v1_xboole_0(X0) => v3_ordinal1(X0)) [input cc3_ordinal1]
19. k1_xboole_0 = o_0_0_xboole_0 [input d2_xboole_0]
32. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
37. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
38. ! [X0] : (v3_ordinal1(X0) => (v3_ordinal1(k1_ordinal1(X0)) & ~v1_xboole_0(k1_ordinal1(X0)))) [input fc2_ordinal1]
46. 1 = k1_ordinal1(k1_xboole_0) [input l14_ordinal3]
55. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
61. ~v1_xboole_0(1) [input spc1_boole]
64. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ~(~r2_hidden(X1,X0) & X0 != X1 & ~r2_hidden(X0,X1)))) [input t14_ordinal1]
68. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_hidden(X0,X1) <=> r1_ordinal1(k1_ordinal1(X0),X1)))) [input t21_ordinal1]
69. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_hidden(X0,k1_ordinal1(X1)) <=> r1_ordinal1(X0,X1)))) [input t22_ordinal1]
71. ! [X0] : r1_tarski(k1_xboole_0,X0) [input t2_xboole_1]
72. ! [X0] : (v3_ordinal1(X0) => k1_xboole_0 = k11_ordinal2(k1_xboole_0,X0)) [input t35_ordinal2]
73. ! [X0] : (v3_ordinal1(X0) => k1_xboole_0 = k11_ordinal2(X0,k1_xboole_0)) [input t38_ordinal2]
74. ! [X0] : (v3_ordinal1(X0) => (k11_ordinal2(X0,1) = X0 & k11_ordinal2(1,X0) = X0)) [input t39_ordinal2]
76. ! [X0] : (r1_tarski(X0,k1_xboole_0) => k1_xboole_0 = X0) [input t3_xboole_1]
77. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r2_hidden(X1,X2) => (r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0))))) [input t40_ordinal2]
79. ! [X0] : ! [X1] : ~(r1_tarski(X1,X0) & r2_hidden(X0,X1)) [input t5_ordinal1]
90. ! [X0,X1] : ~(r1_tarski(X1,X0) & r2_hidden(X0,X1)) [flattening 79]
99. ? [X0] : (? [X1] : (((1 != X1 | 1 != X0) & k11_ordinal2(X0,X1) = 1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
100. ? [X0] : (? [X1] : ((1 != X1 | 1 != X0) & k11_ordinal2(X0,X1) = 1 & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 99]
107. ! [X0] : (v3_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 9]
109. ! [X0] : (k1_xboole_0 = k11_ordinal2(X0,k1_xboole_0) | ~v3_ordinal1(X0)) [ennf transformation 73]
110. ! [X0] : (k1_xboole_0 = k11_ordinal2(k1_xboole_0,X0) | ~v3_ordinal1(X0)) [ennf transformation 72]
112. ! [X0] : ((v3_ordinal1(k1_ordinal1(X0)) & ~v1_xboole_0(k1_ordinal1(X0))) | ~v3_ordinal1(X0)) [ennf transformation 38]
113. ! [X0] : ((k11_ordinal2(X0,1) = X0 & k11_ordinal2(1,X0) = X0) | ~v3_ordinal1(X0)) [ennf transformation 74]
114. ! [X0] : (! [X1] : ((r2_hidden(X0,X1) <=> r1_ordinal1(k1_ordinal1(X0),X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 68]
115. ! [X0] : (! [X1] : ((r2_hidden(X0,k1_ordinal1(X1)) <=> r1_ordinal1(X0,X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 69]
116. ! [X0] : (! [X1] : ((r2_hidden(X1,X0) | X0 = X1 | r2_hidden(X0,X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 64]
117. ! [X0] : (! [X1] : (r2_hidden(X1,X0) | X0 = X1 | r2_hidden(X0,X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 116]
118. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0) | ~r2_hidden(X1,X2)) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 77]
119. ! [X0] : (! [X1] : (! [X2] : (r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0 | ~r2_hidden(X1,X2) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 118]
122. ! [X0] : (k1_xboole_0 = X0 | ~r1_tarski(X0,k1_xboole_0)) [ennf transformation 76]
135. ! [X0,X1] : (~r1_tarski(X1,X0) | ~r2_hidden(X0,X1)) [ennf transformation 90]
157. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 55]
158. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 157]
159. ? [X0] : (? [X1] : ((1 != X1 | 1 != X0) & k11_ordinal2(X0,X1) = 1 & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : ((1 != X1 | 1 != sK0) & 1 = k11_ordinal2(sK0,X1) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
160. ? [X1] : ((1 != X1 | 1 != sK0) & 1 = k11_ordinal2(sK0,X1) & v3_ordinal1(X1)) => ((1 != sK1 | 1 != sK0) & 1 = k11_ordinal2(sK0,sK1) & v3_ordinal1(sK1)) [choice axiom]
161. ((1 != sK1 | 1 != sK0) & 1 = k11_ordinal2(sK0,sK1) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 100,160,159]
162. ! [X0] : (! [X1] : (((r2_hidden(X0,X1) | ~r1_ordinal1(k1_ordinal1(X0),X1)) & (r1_ordinal1(k1_ordinal1(X0),X1) | ~r2_hidden(X0,X1))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [nnf transformation 114]
163. ! [X0] : (! [X1] : (((r2_hidden(X0,k1_ordinal1(X1)) | ~r1_ordinal1(X0,X1)) & (r1_ordinal1(X0,X1) | ~r2_hidden(X0,k1_ordinal1(X1)))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [nnf transformation 115]
170. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 158]
183. v3_ordinal1(sK0) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
184. v3_ordinal1(sK1) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
185. 1 = k11_ordinal2(sK0,sK1) [cnf transformation 161] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
186. 1 != sK1 | 1 != sK0 [cnf transformation 161] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
187. ~v1_xboole_0(1) [cnf transformation 61] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
188. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 32] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
190. v1_xboole_0(k1_xboole_0) [cnf transformation 37] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
192. k1_xboole_0 = o_0_0_xboole_0 [cnf transformation 19] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
195. 1 = k1_ordinal1(k1_xboole_0) [cnf transformation 46] {a:0,w:4,wCS:1600,sine:1,thAx:0,allAx:1,thDist:-1}
207. r1_tarski(k1_xboole_0,X0) [cnf transformation 71] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
215. ~v1_xboole_0(X0) | v3_ordinal1(X0) [cnf transformation 107] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
217. k1_xboole_0 = k11_ordinal2(X0,k1_xboole_0) | ~v3_ordinal1(X0) [cnf transformation 109] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
218. k1_xboole_0 = k11_ordinal2(k1_xboole_0,X0) | ~v3_ordinal1(X0) [cnf transformation 110] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
222. v3_ordinal1(k1_ordinal1(X0)) | ~v3_ordinal1(X0) [cnf transformation 112] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
223. k11_ordinal2(1,X0) = X0 | ~v3_ordinal1(X0) [cnf transformation 113] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
225. r1_ordinal1(k1_ordinal1(X0),X1) | ~r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 162] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
227. ~r2_hidden(X0,k1_ordinal1(X1)) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 163] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
229. r2_hidden(X1,X0) | X0 = X1 | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 117] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
230. r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0 | ~r2_hidden(X1,X2) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 119] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
233. ~r1_tarski(X0,k1_xboole_0) | k1_xboole_0 = X0 [cnf transformation 122] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
244. ~r1_tarski(X1,X0) | ~r2_hidden(X0,X1) [cnf transformation 135] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
264. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 170] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
277. 1 <=> 1 = sK0 [avatar definition]
278. 1 = sK0 <- (1) [avatar component clause 277] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
279. 1 != sK0 <- (~1) [avatar component clause 277] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
281. 2 <=> 1 = sK1 [avatar definition]
283. 1 != sK1 <- (~2) [avatar component clause 281] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
284. ~1 | ~2 [avatar split clause 186,281,277]
286. 3 <=> 1 = k11_ordinal2(sK0,sK1) [avatar definition]
288. 1 = k11_ordinal2(sK0,sK1) <- (3) [avatar component clause 286] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
289. 3 [avatar split clause 185,286]
291. 4 <=> v3_ordinal1(sK1) [avatar definition]
293. v3_ordinal1(sK1) <- (4) [avatar component clause 291] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
294. 4 [avatar split clause 184,291]
296. 5 <=> v3_ordinal1(sK0) [avatar definition]
298. v3_ordinal1(sK0) <- (5) [avatar component clause 296] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
299. 5 [avatar split clause 183,296]
301. 6 <=> v1_xboole_0(1) [avatar definition]
304. ~6 [avatar split clause 187,301]
306. 7 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
309. 7 [avatar split clause 188,306]
316. 9 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
318. v1_xboole_0(k1_xboole_0) <- (9) [avatar component clause 316] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
319. 9 [avatar split clause 190,316]
326. 11 <=> k1_xboole_0 = o_0_0_xboole_0 [avatar definition]
329. 11 [avatar split clause 192,326]
341. 14 <=> 1 = k1_ordinal1(k1_xboole_0) [avatar definition]
343. 1 = k1_ordinal1(k1_xboole_0) <- (14) [avatar component clause 341] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
344. 14 [avatar split clause 195,341]
516. v3_ordinal1(k1_xboole_0) <- (9) [resolution 215,318] {a:1,w:2,wCS:800,sine:1,thAx:0,allAx:2,thDist:-2}
520. 46 <=> v3_ordinal1(k1_xboole_0) [avatar definition]
522. v3_ordinal1(k1_xboole_0) <- (46) [avatar component clause 520] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
523. 46 | ~9 [avatar split clause 516,316,520]
645. v3_ordinal1(1) | ~v3_ordinal1(k1_xboole_0) <- (14) [superposition 222,343] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:2,thDist:-2}
647. 64 <=> v3_ordinal1(1) [avatar definition]
649. v3_ordinal1(1) <- (64) [avatar component clause 647] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
651. ~46 | 64 | ~14 [avatar split clause 645,341,647,520]
710. ~r2_hidden(X0,k1_xboole_0) [resolution 244,207] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
889. r1_ordinal1(1,X1) | ~r2_hidden(k1_xboole_0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(k1_xboole_0) <- (14) [superposition 225,343] {a:1,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
893. ~r2_hidden(k1_xboole_0,X1) | r1_ordinal1(1,X1) | ~v3_ordinal1(X1) <- (14, 46) [subsumption resolution 889,522] {a:1,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
909. ~r2_hidden(X1,1) | r1_ordinal1(X1,k1_xboole_0) | ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(X1) <- (14) [superposition 227,343] {a:1,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
911. ~r2_hidden(X1,1) | r1_ordinal1(X1,k1_xboole_0) | ~v3_ordinal1(X1) <- (14, 46) [subsumption resolution 909,522] {a:1,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
986. 92 <=> 1 = k1_xboole_0 [avatar definition]
987. 1 != k1_xboole_0 <- (~92) [avatar component clause 986] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
998. 94 <=> k1_xboole_0 = sK1 [avatar definition]
999. k1_xboole_0 != sK1 <- (~94) [avatar component clause 998] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1000. k1_xboole_0 = sK1 <- (94) [avatar component clause 998] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1015. 97 <=> ! [X8] : (r2_hidden(k11_ordinal2(X8,sK1),1) | ~v3_ordinal1(X8) | ~r2_hidden(X8,sK0)) [avatar definition]
1016. r2_hidden(k11_ordinal2(X8,sK1),1) | ~v3_ordinal1(X8) | ~r2_hidden(X8,sK0) <- (97) [avatar component clause 1015] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1018. k1_xboole_0 != o_0_0_xboole_0 | 1 != k1_xboole_0 | v1_xboole_0(1) | ~v1_xboole_0(o_0_0_xboole_0) [theory tautology sat conflict] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:1,thDist:-1}
1021. 1 = k11_ordinal2(sK0,k1_xboole_0) <- (3, 94) [backward demodulation 288,1000] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1024. 98 <=> 1 = k11_ordinal2(sK0,k1_xboole_0) [avatar definition]
1026. 1 = k11_ordinal2(sK0,k1_xboole_0) <- (98) [avatar component clause 1024] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1027. 98 | ~3 | ~94 [avatar split clause 1021,998,286,1024]
1051. 99 <=> r2_hidden(sK0,1) [avatar definition]
1053. ~r2_hidden(sK0,1) <- (~99) [avatar component clause 1051] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1086. 104 <=> r2_hidden(1,sK0) [avatar definition]
1090. 105 <=> r2_hidden(sK1,1) [avatar definition]
1092. r2_hidden(sK1,1) <- (105) [avatar component clause 1090] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1097. 106 <=> r2_hidden(k1_xboole_0,sK0) [avatar definition]
1098. r2_hidden(k1_xboole_0,sK0) <- (106) [avatar component clause 1097] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1099. ~r2_hidden(k1_xboole_0,sK0) <- (~106) [avatar component clause 1097] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1116. 1 = k11_ordinal2(1,sK1) <- (1, 3) [backward demodulation 288,278] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1124. 108 <=> 1 = k11_ordinal2(1,sK1) [avatar definition]
1126. 1 = k11_ordinal2(1,sK1) <- (108) [avatar component clause 1124] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1127. 108 | ~1 | ~3 [avatar split clause 1116,286,277,1124]
1149. 1 = k1_xboole_0 | ~v3_ordinal1(sK0) <- (98) [superposition 217,1026] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1155. 1 = k1_xboole_0 <- (5, 98) [subsumption resolution 1149,298] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1156. 92 | ~5 | ~98 [avatar split clause 1155,1024,296,986]
1350. 123 <=> k1_xboole_0 = sK0 [avatar definition]
1352. k1_xboole_0 = sK0 <- (123) [avatar component clause 1350] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1395. k1_xboole_0 = sK0 | r2_hidden(sK0,k1_xboole_0) | ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(sK0) <- (~106) [resolution 1099,229] {a:3,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1397. k1_xboole_0 = sK0 | ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(sK0) <- (~106) [subsumption resolution 1395,710] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1398. k1_xboole_0 = sK0 | ~v3_ordinal1(sK0) <- (46, ~106) [subsumption resolution 1397,522] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1399. k1_xboole_0 = sK0 <- (5, 46, ~106) [subsumption resolution 1398,298] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1400. 123 | ~5 | ~46 | 106 [avatar split clause 1399,1097,520,296,1350]
1414. 1 = sK1 | ~v3_ordinal1(sK1) <- (108) [superposition 223,1126] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1423. ~v3_ordinal1(sK1) <- (~2, 108) [subsumption resolution 1414,283] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1424. $false <- (~2, 4, 108) [subsumption resolution 1423,293] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1425. 2 | ~4 | ~108 [avatar contradiction clause 1424]
1534. 1 = k11_ordinal2(k1_xboole_0,sK1) <- (3, 123) [backward demodulation 288,1352] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1552. 133 <=> 1 = k11_ordinal2(k1_xboole_0,sK1) [avatar definition]
1554. 1 = k11_ordinal2(k1_xboole_0,sK1) <- (133) [avatar component clause 1552] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1555. 133 | ~3 | ~123 [avatar split clause 1534,1350,286,1552]
1569. 1 = k1_xboole_0 | ~v3_ordinal1(sK1) <- (133) [superposition 218,1554] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1576. ~v3_ordinal1(sK1) <- (~92, 133) [subsumption resolution 1569,987] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1577. $false <- (4, ~92, 133) [subsumption resolution 1576,293] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1578. ~4 | 92 | ~133 [avatar contradiction clause 1577]
1586. r2_hidden(k11_ordinal2(X0,sK1),1) | k1_xboole_0 = sK1 | ~r2_hidden(X0,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK1) <- (3) [superposition 230,288] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1589. r2_hidden(k11_ordinal2(X0,sK1),1) | ~r2_hidden(X0,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK1) <- (3, ~94) [subsumption resolution 1586,999] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1590. r2_hidden(k11_ordinal2(X0,sK1),1) | ~r2_hidden(X0,sK0) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK1) <- (3, 5, ~94) [subsumption resolution 1589,298] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1591. r2_hidden(k11_ordinal2(X0,sK1),1) | ~r2_hidden(X0,sK0) | ~v3_ordinal1(X0) <- (3, 4, 5, ~94) [subsumption resolution 1590,293] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1592. 97 | ~3 | ~4 | ~5 | 94 [avatar split clause 1591,998,296,291,286,1015]
1634. r2_hidden(sK1,1) | ~v3_ordinal1(1) | ~r2_hidden(1,sK0) | ~v3_ordinal1(sK1) <- (97) [superposition 1016,223] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1738. r1_ordinal1(1,sK0) | ~v3_ordinal1(sK0) <- (14, 46, 106) [resolution 893,1098] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1746. r1_ordinal1(1,sK0) <- (5, 14, 46, 106) [subsumption resolution 1738,298] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1748. 137 <=> r1_ordinal1(1,sK0) [avatar definition]
1750. r1_ordinal1(1,sK0) <- (137) [avatar component clause 1748] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1751. 137 | ~5 | ~14 | ~46 | ~106 [avatar split clause 1746,1097,520,341,296,1748]
1767. r1_tarski(1,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(1) <- (137) [resolution 1750,264] {a:4,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1769. r1_tarski(1,sK0) | ~v3_ordinal1(1) <- (5, 137) [subsumption resolution 1767,298] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1770. r1_tarski(1,sK0) <- (5, 64, 137) [subsumption resolution 1769,649] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1772. 138 <=> r1_tarski(1,sK0) [avatar definition]
1774. r1_tarski(1,sK0) <- (138) [avatar component clause 1772] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1775. 138 | ~5 | ~64 | ~137 [avatar split clause 1770,1748,647,296,1772]
1790. ~r2_hidden(sK0,1) <- (138) [resolution 1774,244] {a:5,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1801. ~99 | ~138 [avatar split clause 1790,1772,1051]
1811. r2_hidden(sK1,1) | ~r2_hidden(1,sK0) | ~v3_ordinal1(sK1) <- (64, 97) [subsumption resolution 1634,649] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1812. r2_hidden(sK1,1) | ~r2_hidden(1,sK0) <- (4, 64, 97) [subsumption resolution 1811,293] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1813. ~104 | 105 | ~4 | ~64 | ~97 [avatar split clause 1812,1015,647,291,1090,1086]
1870. 1 = sK0 | r2_hidden(1,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(1) <- (~99) [resolution 1053,229] {a:3,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1871. r2_hidden(1,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(1) <- (~1, ~99) [subsumption resolution 1870,279] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1872. r2_hidden(1,sK0) | ~v3_ordinal1(1) <- (~1, 5, ~99) [subsumption resolution 1871,298] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1873. r2_hidden(1,sK0) <- (~1, 5, 64, ~99) [subsumption resolution 1872,649] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1874. 104 | 1 | ~5 | ~64 | 99 [avatar split clause 1873,1051,647,296,277,1086]
1971. r1_ordinal1(sK1,k1_xboole_0) | ~v3_ordinal1(sK1) <- (14, 46, 105) [resolution 1092,911] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1976. r1_ordinal1(sK1,k1_xboole_0) <- (4, 14, 46, 105) [subsumption resolution 1971,293] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1978. 146 <=> r1_ordinal1(sK1,k1_xboole_0) [avatar definition]
1980. r1_ordinal1(sK1,k1_xboole_0) <- (146) [avatar component clause 1978] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1981. 146 | ~4 | ~14 | ~46 | ~105 [avatar split clause 1976,1090,520,341,291,1978]
2114. r1_tarski(sK1,k1_xboole_0) | ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(sK1) <- (146) [resolution 1980,264] {a:4,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
2116. r1_tarski(sK1,k1_xboole_0) | ~v3_ordinal1(sK1) <- (46, 146) [subsumption resolution 2114,522] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
2117. r1_tarski(sK1,k1_xboole_0) <- (4, 46, 146) [subsumption resolution 2116,293] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
2119. 149 <=> r1_tarski(sK1,k1_xboole_0) [avatar definition]
2121. r1_tarski(sK1,k1_xboole_0) <- (149) [avatar component clause 2119] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
2122. 149 | ~4 | ~46 | ~146 [avatar split clause 2117,1978,520,291,2119]
2237. k1_xboole_0 = sK1 <- (149) [resolution 2121,233] {a:5,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
2249. $false <- (~94, 149) [subsumption resolution 2237,999] {a:5,w:0,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
2250. 94 | ~149 [avatar contradiction clause 2249]
2257. $false [avatar sat refutation 284,289,294,299,304,309,319,329,344,523,651,1018,1027,1127,1156,1400,1425,1555,1578,1592,1751,1775,1801,1813,1874,1981,2122,2250] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1451,thDist:-1451}
% SZS output end Proof for ordinal3__t37_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 82
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 93
% Generated clauses: 1464
% Active clauses: 284
% Passive clauses: 515
% Final active clauses: 226
% Final passive clauses: 200
% Main loop iterations started: 299
% 
% >>> Simplifying Inferences
% Duplicate literals: 103
% Trivial inequalities: 5
% Fw subsumption resolutions: 252
% Fw demodulations: 11
% Bw demodulations: 28
% 
% >>> Deletion Inferences
% Simple tautologies: 39
% Equational tautologies: 9
% Forward subsumptions: 327
% 
% >>> Generating Inferences
% Binary resolution: 487
% Unit resulting resolution: 5
% Factoring: 2
% Forward superposition: 178
% Backward superposition: 95
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 33
% Split components: 66
% Sat splitting refutations: 17
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 224
% SAT solver unit clauses: 33
% SAT solver binary clauses: 93
% 
% Memory used [KB]: 6268
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% sat proof minimization: 0.001 s
% congruence closure: 0.002 s
% unit clause index maintenance: 0.001 s
% forward subsumption: 0.004 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.003 s
% other: 0.015 s ( own 0.002 s ) 

