i: [1,101,0,6,1,1,0,0] t46_setfam_1
a: [3,148,0,6,1,101]
s: 148
i: [1,102,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,153,0,2,1,102]
s: 153
i: [1,103,0,2,1,0,0,1] fc1_xboole_0
a: [3,158,0,2,1,103]
s: 158
i: [1,140,0,2,1,0,0,3] rc2_setfam_1
a: [3,168,0,2,1,140]
s: 168
i: [1,141,0,2,1,0,0,2] rc2_xboole_0
a: [3,173,0,2,1,141]
s: 173
i: [1,142,0,2,1,0,0,2] rc1_xboole_0
a: [3,178,0,2,1,142]
s: 178
i: [1,144,0,2,1,0,0,2] rc1_setfam_1
a: [3,183,0,2,1,144]
s: 183
i: [1,143,0,2,1,0,0,2] rc1_setfam_1
a: [3,188,0,2,1,143]
s: 188
i: [1,105,0,3,1,0,0,1] fc2_xboole_0
s: 105
i: [1,106,0,3,1,0,0,1] fc1_subset_1
s: 106
i: [1,107,0,3,1,0,0,2] t2_xboole_1
s: 107
i: [1,135,0,3,1,0,0,1] rc2_subset_1
s: 135
i: [1,138,0,3,1,0,0,2] reflexivity_r1_tarski
s: 138
i: [1,104,0,3,1,0,0,255] d2_xboole_0
a: [3,163,0,3,1,104]
s: 163
r: 153
i: [1,108,0,4,2,0,0,3] cc2_setfam_1
s: 108
i: [1,133,0,4,1,0,0,1] existence_m1_subset_1
s: 133
d: [2,191,1,2,1,1,40,108,158]
a: [3,197,1,2,1,191]
s: 197
i: [1,137,0,4,1,0,0,255] rc3_subset_1
s: 137
d: [2,193,1,2,1,1,40,108,178]
a: [3,202,1,2,1,193]
s: 202
i: [1,110,0,5,2,0,0,3] rc3_setfam_1
s: 110
i: [1,111,0,5,2,0,0,1] fc2_setfam_1
s: 111
d: [2,192,1,3,1,0,40,108,135]
s: 192
i: [1,115,0,5,2,0,0,1] rc1_subset_1
s: 115
i: [1,118,0,5,2,0,0,255] t6_boole
s: 118
i: [1,130,0,5,2,0,0,2] t7_boole
s: 130
i: [1,134,0,5,1,0,0,1] rc2_subset_1
s: 134
d: [2,208,1,8,2,1,44,148,118]
s: 208
i: [1,136,0,5,1,0,0,255] rc3_subset_1
s: 136
d: [2,206,1,5,2,0,44,107,118]
s: 206
i: [1,109,0,6,2,0,0,3] rc3_setfam_1
s: 109
i: [1,117,0,6,2,0,0,255] rc4_subset_1
s: 117
i: [1,121,0,6,3,0,0,2] cc3_setfam_1
s: 121
i: [1,123,0,6,2,0,0,2] t1_subset
s: 123
i: [1,139,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 139
i: [1,114,0,7,2,0,0,1] rc1_subset_1
s: 114
i: [1,116,0,7,2,0,0,255] rc4_subset_1
s: 116
i: [1,125,0,7,2,0,0,1] t31_zfmisc_1
s: 125
i: [1,126,0,7,2,0,0,1] t31_zfmisc_1
s: 126
d: [2,212,1,4,1,0,40,125,138]
s: 212
i: [1,127,0,7,2,0,0,2] t3_subset
s: 127
i: [1,128,0,7,2,0,0,2] t3_subset
s: 128
d: [2,224,1,5,1,1,40,128,148]
a: [3,230,1,5,1,224]
s: 230
i: [1,129,0,7,3,0,0,255] t8_boole
s: 129
d: [2,232,2,4,1,1,40,230,126]
a: [3,238,2,4,1,232]
s: 238
i: [1,119,0,8,3,0,0,1] cc1_subset_1
s: 119
i: [1,124,0,8,3,0,0,2] t2_subset
s: 124
d: [2,240,3,6,2,1,44,238,118]
s: 240
i: [1,120,0,9,3,0,0,255] cc4_subset_1
s: 120
d: [2,225,2,7,2,1,40,128,208]
s: 225
i: [1,122,0,9,4,0,0,2] cc1_setfam_1
s: 122
i: [1,132,0,9,3,0,0,2] t5_subset
s: 132
d: [2,234,2,7,2,1,44,230,118]
s: 234
i: [1,131,0,10,3,0,0,2] t4_subset
s: 131
d: [2,271,1,11,4,0,40,122,128]
d: [2,278,1,8,3,0,21,271,106]
d: [2,279,1,5,2,0,21,278,206]
a: [3,282,1,3,1,279]
s: 282
i: [1,112,0,11,4,0,0,255] cc3_subset_1
s: 112
i: [1,113,0,11,4,0,0,255] cc2_subset_1
s: 113
d: [2,222,1,4,1,0,40,127,134]
s: 222
d: [2,223,1,4,1,0,40,127,136]
s: 223
d: [2,215,2,4,1,0,40,212,139]
s: 215
d: [2,221,1,5,1,0,40,127,133]
s: 221
d: [2,247,1,5,2,0,40,119,136]
s: 247
d: [2,216,2,4,1,0,40,212,123]
s: 216
d: [2,242,1,9,3,0,40,119,109]
d: [2,248,1,5,2,0,21,242,110]
s: 248
d: [2,254,1,8,2,0,40,124,134]
d: [2,260,1,5,1,0,21,254,106]
s: 260
d: [2,323,4,0,0,1,52,135,260,240]
e: 323
d: [2,324,4,3,1,1,40,260,240]
d: [2,329,4,0,0,1,21,324,135]
e: 329
f: 329,232,224,101
% Refutation found. Thanks to Tanya!
% SZS status Theorem for setfam_1__t46_setfam_1
% SZS output start Proof for setfam_1__t46_setfam_1
1. ! [X0] : m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(X0))) [input t46_setfam_1]
2. ~! [X0] : m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(X0))) [negated conjecture 1]
18. ! [X0] : ~v1_xboole_0(k1_zfmisc_1(X0)) [input fc1_subset_1]
26. ! [X0] : ? [X1] : (v1_xboole_0(X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) [input rc2_subset_1]
33. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
35. ! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [input t31_zfmisc_1]
36. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
39. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
43. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 33]
44. ! [X0,X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [flattening 35]
45. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 36]
51. ? [X0] : ~m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(X0))) [ennf transformation 2]
61. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 39]
69. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 43]
70. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 69]
77. ? [X0] : ~m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(X0))) => ~m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(sK0))) [choice axiom]
78. ~m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(sK0))) [skolemisation 51,77]
85. ! [X0,X1] : ((r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X1) | ~r1_tarski(k1_tarski(X0),X1))) [nnf transformation 44]
86. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 45]
89. ! [X0] : (? [X1] : (v1_xboole_0(X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (v1_xboole_0(sK5(X0)) & m1_subset_1(sK5(X0),k1_zfmisc_1(X0)))) [choice axiom]
90. ! [X0] : (v1_xboole_0(sK5(X0)) & m1_subset_1(sK5(X0),k1_zfmisc_1(X0))) [skolemisation 26,89]
101. ~m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(sK0))) [cnf transformation 78] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
106. ~v1_xboole_0(k1_zfmisc_1(X0)) [cnf transformation 18] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
118. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 61] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
124. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 70] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
126. r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 85] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
128. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 86] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
134. m1_subset_1(sK5(X0),k1_zfmisc_1(X0)) [cnf transformation 90] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
135. v1_xboole_0(sK5(X0)) [cnf transformation 90] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
146. 1 <=> m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(sK0))) [avatar definition]
148. ~m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(sK0))) <- (~1) [avatar component clause 146] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
149. ~1 [avatar split clause 101,146]
224. ~r1_tarski(k1_tarski(k1_xboole_0),k1_zfmisc_1(sK0)) <- (~1) [resolution 128,148] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
228. 12 <=> r1_tarski(k1_tarski(k1_xboole_0),k1_zfmisc_1(sK0)) [avatar definition]
230. ~r1_tarski(k1_tarski(k1_xboole_0),k1_zfmisc_1(sK0)) <- (~12) [avatar component clause 228] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
231. ~12 | 1 [avatar split clause 224,146,228]
232. ~r2_hidden(k1_xboole_0,k1_zfmisc_1(sK0)) <- (~12) [resolution 230,126] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
236. 13 <=> r2_hidden(k1_xboole_0,k1_zfmisc_1(sK0)) [avatar definition]
238. ~r2_hidden(k1_xboole_0,k1_zfmisc_1(sK0)) <- (~13) [avatar component clause 236] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
239. ~13 | 12 [avatar split clause 232,228,236]
240. ~r2_hidden(X0,k1_zfmisc_1(sK0)) | ~v1_xboole_0(X0) <- (~13) [superposition 238,118] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
254. v1_xboole_0(k1_zfmisc_1(X6)) | r2_hidden(sK5(X6),k1_zfmisc_1(X6)) [resolution 124,134] {a:1,w:8,wCS:3200,sine:1,thAx:0,allAx:3,thDist:-3}
260. r2_hidden(sK5(X6),k1_zfmisc_1(X6)) [subsumption resolution 254,106] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
324. ~v1_xboole_0(sK5(sK0)) <- (~13) [resolution 260,240] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
329. $false <- (~13) [subsumption resolution 324,135] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
330. 13 [avatar contradiction clause 329]
331. $false [avatar sat refutation 149,231,239,330] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:70,thDist:-70}
% SZS output end Proof for setfam_1__t46_setfam_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 40
% Introduced skolems: 11
% 
% >>> Saturation
% Initial clauses: 44
% Generated clauses: 169
% Active clauses: 62
% Passive clauses: 101
% Final active clauses: 61
% Final passive clauses: 37
% Main loop iterations started: 64
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 18
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 82
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 19
% SAT solver unit clauses: 12
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

