i: [1,198,0,5,2,1,0,0] l57_mod_2
a: [3,294,0,2,1,198]
s: 294
i: [1,177,0,2,1,1,0,0] l57_mod_2
a: [3,331,0,2,1,177]
s: 331
i: [1,176,0,2,1,1,0,0] l57_mod_2
a: [3,336,0,2,1,176]
s: 336
i: [1,175,0,2,1,1,0,0] l57_mod_2
a: [3,341,0,2,1,175]
s: 341
i: [1,202,0,6,2,1,0,0] l57_mod_2
a: [3,270,0,3,1,202]
s: 270
i: [1,201,0,6,2,1,0,0] l57_mod_2
a: [3,279,0,3,1,201]
s: 279
i: [1,199,0,9,2,1,0,0] l57_mod_2
a: [3,289,0,6,1,199]
s: 289
i: [1,200,0,11,2,1,0,0] l57_mod_2
a: [3,284,0,8,1,200]
s: 284
i: [1,203,0,2,1,0,0,255] fc1_xboole_0
a: [3,346,0,2,1,203]
s: 346
i: [1,261,0,2,1,0,0,1] rc2_xboole_0
a: [3,351,0,2,1,261]
s: 351
i: [1,262,0,2,1,0,0,1] rc1_xboole_0
a: [3,356,0,2,1,262]
s: 356
i: [1,264,0,2,1,0,0,1] rc1_classes2
a: [3,361,0,2,1,264]
s: 361
i: [1,263,0,2,1,0,0,1] rc1_classes2
a: [3,366,0,2,1,263]
s: 366
i: [1,204,0,3,1,0,0,1] fc1_subset_1
s: 204
i: [1,236,0,3,1,0,0,1] rc2_subset_1
s: 236
i: [1,239,0,3,1,0,0,255] reflexivity_r1_tarski
s: 239
i: [1,205,0,4,1,0,0,2] fc2_funct_2
s: 205
i: [1,213,0,4,2,0,0,255] cc2_classes2
s: 213
d: [2,368,1,2,1,1,40,213,331]
a: [3,373,1,2,1,368]
s: 373
i: [1,214,0,4,2,0,0,255] cc2_classes2
s: 214
i: [1,234,0,4,1,0,0,1] existence_m1_subset_1
s: 234
d: [2,380,1,2,1,1,40,214,331]
a: [3,385,1,2,1,380]
s: 385
i: [1,238,0,4,1,0,0,255] rc3_subset_1
s: 238
d: [2,369,1,2,1,1,40,213,361]
a: [3,378,1,2,1,369]
s: 378
i: [1,257,0,4,1,0,0,255] rc1_funct_2
s: 257
i: [1,260,0,4,1,0,0,255] rc1_relset_1
s: 260
d: [2,381,1,2,1,1,40,214,361]
a: [3,390,1,2,1,381]
s: 390
i: [1,210,0,5,2,0,0,1] rc1_subset_1
s: 210
i: [1,215,0,5,2,0,0,255] t6_boole
s: 215
i: [1,226,0,5,2,0,0,1] t7_boole
s: 226
i: [1,235,0,5,1,0,0,1] rc2_subset_1
s: 235
i: [1,237,0,5,1,0,0,255] rc3_subset_1
s: 237
i: [1,212,0,6,2,0,0,255] rc4_subset_1
s: 212
i: [1,220,0,6,3,0,0,255] cc3_classes2
s: 220
i: [1,221,0,6,2,0,0,1] t1_subset
s: 221
i: [1,240,0,6,2,0,0,2] fc1_funct_2
s: 240
d: [2,397,1,3,1,1,40,221,279]
a: [3,401,1,3,1,397]
s: 401
i: [1,247,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 247
d: [2,403,1,3,1,1,40,247,279]
a: [3,407,1,3,1,403]
s: 407
i: [1,253,0,6,2,0,0,255] dt_m1_funct_2
s: 253
i: [1,255,0,6,1,0,0,255] existence_m1_funct_2
s: 255
d: [2,409,1,4,1,0,40,255,253]
s: 409
i: [1,258,0,6,1,0,0,255] rc1_funct_2
s: 258
i: [1,209,0,7,2,0,0,1] rc1_subset_1
s: 209
i: [1,211,0,7,2,0,0,255] rc4_subset_1
s: 211
i: [1,219,0,7,3,0,0,1] rc2_classes2
s: 219
i: [1,224,0,7,2,0,0,255] t3_subset
s: 224
i: [1,225,0,7,3,0,0,255] t8_boole
s: 225
d: [2,392,1,7,3,0,44,215,215]
s: 392
i: [1,216,0,8,3,0,0,1] cc1_subset_1
s: 216
i: [1,218,0,8,3,0,0,1] rc2_classes2
s: 218
d: [2,411,1,8,2,1,40,216,284]
a: [3,420,1,6,1,411]
s: 420
i: [1,222,0,8,3,0,0,1] t2_subset
s: 222
d: [2,430,1,15,2,1,40,222,284]
d: [2,439,1,8,1,1,21,430,204]
a: [3,443,1,8,1,439]
s: 443
i: [1,241,0,8,2,0,0,255] dt_k9_funct_2
s: 241
i: [1,248,0,8,3,0,0,1] fc10_subset_1
s: 248
d: [2,449,2,8,1,1,40,443,247]
a: [3,455,2,8,1,449]
s: 455
i: [1,252,0,8,3,0,0,2] fc3_funct_2
s: 252
d: [2,416,1,5,2,0,40,216,237]
s: 416
i: [1,256,0,8,1,0,0,255] rc1_funct_2
s: 256
i: [1,259,0,8,1,0,0,255] rc1_relset_1
s: 259
d: [2,435,1,8,2,0,40,222,235]
d: [2,447,1,5,1,0,21,435,204]
s: 447
i: [1,217,0,9,3,0,0,255] cc4_subset_1
s: 217
d: [2,436,1,8,2,0,40,222,237]
d: [2,448,1,5,1,0,21,436,204]
s: 448
i: [1,233,0,9,3,0,0,1] t5_subset
s: 233
i: [1,242,0,9,2,0,0,255] redefinition_k9_funct_2
s: 242
d: [2,466,2,5,1,0,40,447,247]
s: 466
i: [1,232,0,10,3,0,0,1] t4_subset
s: 232
d: [2,499,1,10,2,1,40,232,284]
s: 499
i: [1,244,0,10,3,0,0,1] cc4_relset_1
s: 244
i: [1,245,0,10,3,0,0,1] cc3_relset_1
s: 245
d: [2,524,1,6,2,1,40,245,284]
a: [3,536,1,4,1,524]
s: 536
i: [1,207,0,11,4,0,0,255] cc3_subset_1
s: 207
a: [3,423,1,2,1,411]
s: 423
i: [1,208,0,11,4,0,0,255] cc2_subset_1
s: 208
i: [1,227,0,11,4,0,0,255] t10_ordinal1
s: 227
d: [2,561,1,8,3,1,40,227,279]
d: [2,565,1,6,2,2,21,561,373]
s: 565
i: [1,223,0,13,3,0,0,255] cc4_funct_2
s: 223
d: [2,562,2,18,3,1,40,227,443]
a: [3,569,2,7,1,562]
s: 569
i: [1,254,0,13,3,0,0,255] cc1_funct_2
s: 254
i: [1,228,0,15,5,0,0,1] t61_classes2
s: 228
d: [2,509,2,16,3,1,40,499,222]
d: [2,510,2,10,2,2,21,509,420]
s: 510
d: [2,596,0,13,4,0,21,228,226]
s: 596
d: [2,613,2,13,4,2,40,596,565]
a: [3,618,2,2,1,613]
s: 618
i: [1,229,0,15,5,0,0,1] t61_classes2
s: 229
d: [2,623,0,13,4,0,21,229,226]
s: 623
d: [2,597,3,19,4,2,40,510,227]
a: [3,604,3,6,1,597]
s: 604
i: [1,243,0,15,4,0,0,255] cc3_funct_2
s: 243
d: [2,632,1,13,3,1,40,243,284]
d: [2,641,1,7,2,2,21,632,289]
d: [2,642,1,5,1,3,21,641,341]
a: [3,646,1,5,1,642]
s: 646
i: [1,246,0,15,4,0,0,255] cc2_funct_2
s: 246
i: [1,265,0,17,4,0,0,255] t8_funct_2
s: 265
d: [2,598,3,10,2,2,40,510,247]
s: 598
i: [1,250,0,18,6,0,0,1] cc8_funct_2
s: 250
d: [2,676,4,14,2,2,40,598,510]
d: [2,677,4,7,1,2,18,676]
a: [3,682,4,7,1,677]
s: 682
i: [1,230,0,20,5,0,0,255] t8_funct_2
s: 230
r: 294
a: [3,293,0,2,1,198]
s: 293
i: [1,182,0,5,2,1,0,0] l57_mod_2
a: [3,322,0,2,1,182]
s: 322
a: [3,274,0,3,1,202]
s: 274
d: [2,699,5,10,3,1,40,682,596]
d: [2,700,5,5,2,1,21,699,596]
a: [3,704,5,2,1,700]
s: 704
i: [1,187,0,7,2,1,0,0] l57_mod_2
a: [3,313,0,4,1,187]
s: 313
i: [1,192,0,9,2,1,0,0] l57_mod_2
a: [3,304,0,6,1,192]
s: 304
d: [2,740,1,12,5,1,40,304,250]
d: [2,749,1,10,4,1,18,740]
d: [2,764,1,6,3,2,21,749,313]
d: [2,765,1,4,2,3,21,764,322]
d: [2,766,1,2,1,4,21,765,341]
a: [3,770,1,2,1,766]
s: 770
i: [1,206,0,29,10,0,0,1] t28_mod_2
s: 206
d: [2,736,1,9,3,1,40,304,243]
d: [2,757,1,5,2,2,21,736,313]
d: [2,758,1,3,1,3,21,757,341]
a: [3,762,1,3,1,758]
s: 762
d: [2,786,0,27,9,0,21,206,226]
s: 786
d: [2,787,1,0,0,7,52,341,341,331,322,279,274,279,313,304,786]
e: 787
r: 293
r: 313
r: 762
s: 294
d: [2,711,1,18,4,1,40,230,284]
d: [2,720,1,12,3,2,21,711,289]
d: [2,721,1,10,2,3,21,720,294]
a: [3,728,1,3,1,721]
s: 728
a: [3,312,0,4,1,187]
s: 312
a: [3,761,1,3,1,758]
s: 761
d: [2,735,1,14,4,1,40,304,230]
d: [2,750,1,10,3,2,21,735,313]
d: [2,751,1,8,2,3,21,750,322]
a: [3,755,1,5,1,751]
s: 755
d: [2,748,1,11,2,1,40,304,222]
d: [2,780,1,6,1,1,21,748,204]
a: [3,784,1,6,1,780]
s: 784
d: [2,832,2,12,3,1,40,755,227]
a: [3,840,2,4,1,832]
s: 840
a: [3,725,1,7,1,721]
s: 725
d: [2,878,3,6,2,1,44,840,242]
d: [2,879,3,4,1,2,21,878,341]
a: [3,883,3,4,1,879]
s: 883
d: [2,744,1,8,2,1,40,304,232]
s: 744
d: [2,790,1,27,8,1,40,786,284]
d: [2,810,1,21,7,2,21,790,289]
d: [2,811,1,19,6,3,21,810,341]
d: [2,812,1,15,5,4,21,811,536]
d: [2,813,1,10,4,4,21,812,596]
a: [3,816,1,8,3,813]
s: 816
d: [2,921,2,0,0,4,52,331,270,279,816]
e: 921
d: [2,922,2,5,2,2,40,816,279]
d: [2,925,2,3,1,3,21,922,331]
d: [2,926,2,0,0,4,21,925,270]
e: 926
r: 598
r: 510
r: 420
r: 423
r: 536
r: 816
r: 770
a: [3,424,1,2,1,411]
s: 424
a: [3,769,1,2,1,766]
s: 769
d: [2,484,1,9,2,1,40,233,284]
a: [3,935,1,3,1,484]
s: 935
r: 455
r: 682
r: 499
d: [2,745,1,7,2,1,40,304,233]
a: [3,954,1,3,1,745]
s: 954
r: 744
a: [3,535,1,4,1,524]
s: 535
d: [2,962,2,0,0,2,52,341,341,535,248]
e: 962
d: [2,963,2,4,2,1,40,535,248]
d: [2,964,2,2,1,1,18,963]
d: [2,966,2,0,0,2,21,964,341]
e: 966
i: [1,196,0,6,2,1,0,0] l57_mod_2
d: [2,266,0,3,1,0,18,196]
i: [1,190,0,14,2,1,0,0] l57_mod_2
i: [1,189,0,12,2,1,0,0] l57_mod_2
i: [1,188,0,8,2,1,0,0] l57_mod_2
i: [1,185,0,12,2,1,0,0] l57_mod_2
i: [1,184,0,10,2,1,0,0] l57_mod_2
i: [1,183,0,6,2,1,0,0] l57_mod_2
i: [1,180,0,10,2,1,0,0] l57_mod_2
i: [1,179,0,8,2,1,0,0] l57_mod_2
i: [1,178,0,4,2,1,0,0] l57_mod_2
d: [2,946,1,17,5,4,21,811,294]
d: [2,947,1,12,4,4,21,946,596]
f: 966,947,926,787,175,177,178,179,180,182,183,184,185,187,188,189,190,192,266,198,199,200,202
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mod_2__l57_mod_2
% SZS output start Proof for mod_2__l57_mod_2
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((v1_classes2(X1) & ~v1_xboole_0(X1)) => (! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_funct_2(X2,X0,X0) & v1_funct_1(X2)) => (r2_hidden(X0,X1) => r2_hidden(X2,X1))) & ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X0),X0))) & v1_funct_2(X2,k2_zfmisc_1(X0,X0),X0) & v1_funct_1(X2)) => (r2_hidden(X0,X1) => r2_hidden(X2,X1)))))) [input l57_mod_2]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((v1_classes2(X1) & ~v1_xboole_0(X1)) => (! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_funct_2(X2,X0,X0) & v1_funct_1(X2)) => (r2_hidden(X0,X1) => r2_hidden(X2,X1))) & ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X0),X0))) & v1_funct_2(X2,k2_zfmisc_1(X0,X0),X0) & v1_funct_1(X2)) => (r2_hidden(X0,X1) => r2_hidden(X2,X1)))))) [negated conjecture 1]
31. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_zfmisc_1(X0,X1))) [input fc10_subset_1]
51. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : ((v1_classes2(X2) & ~v1_xboole_0(X2)) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_2(X3,X0,X1) & v1_funct_1(X3)) => ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) => r2_hidden(X3,X2)))))) [input t28_mod_2]
56. ! [X0] : ! [X1] : ! [X2] : ((v1_classes2(X2) & ~v1_xboole_0(X2)) => ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) => (r2_hidden(k1_funct_2(X0,X1),X2) & r2_hidden(k2_zfmisc_1(X0,X1),X2)))) [input t61_classes2]
58. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
61. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((v1_classes2(X1) & ~v1_xboole_0(X1)) => (! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_funct_2(X2,X0,X0) & v1_funct_1(X2)) => (r2_hidden(X0,X1) => r2_hidden(X2,X1))) & ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X0),X0))) & v1_funct_2(X3,k2_zfmisc_1(X0,X0),X0) & v1_funct_1(X3)) => (r2_hidden(X0,X1) => r2_hidden(X3,X1)))))) [rectify 2]
67. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 58]
69. ! [X0,X1,X2] : ((v1_classes2(X2) & ~v1_xboole_0(X2)) => ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) => (r2_hidden(k1_funct_2(X0,X1),X2) & r2_hidden(k2_zfmisc_1(X0,X1),X2)))) [flattening 56]
92. ? [X0] : (? [X1] : ((? [X2] : ((~r2_hidden(X2,X1) & r2_hidden(X0,X1)) & (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_funct_2(X2,X0,X0) & v1_funct_1(X2))) | ? [X3] : ((~r2_hidden(X3,X1) & r2_hidden(X0,X1)) & (m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X0),X0))) & v1_funct_2(X3,k2_zfmisc_1(X0,X0),X0) & v1_funct_1(X3)))) & (v1_classes2(X1) & ~v1_xboole_0(X1))) & ~v1_xboole_0(X0)) [ennf transformation 61]
93. ? [X0] : (? [X1] : ((? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X0,X1) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_funct_2(X2,X0,X0) & v1_funct_1(X2)) | ? [X3] : (~r2_hidden(X3,X1) & r2_hidden(X0,X1) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X0),X0))) & v1_funct_2(X3,k2_zfmisc_1(X0,X0),X0) & v1_funct_1(X3))) & v1_classes2(X1) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [flattening 92]
94. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X1,X2) | ~r2_hidden(X0,X2))) | (~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~v1_funct_2(X3,X0,X1) | ~v1_funct_1(X3))) | (~v1_classes2(X2) | v1_xboole_0(X2))) | v1_xboole_0(X1)) | v1_xboole_0(X0)) [ennf transformation 51]
95. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (r2_hidden(X3,X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~v1_funct_2(X3,X0,X1) | ~v1_funct_1(X3)) | ~v1_classes2(X2) | v1_xboole_0(X2)) | v1_xboole_0(X1)) | v1_xboole_0(X0)) [flattening 94]
117. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 67]
120. ! [X0,X1,X2] : (((r2_hidden(k1_funct_2(X0,X1),X2) & r2_hidden(k2_zfmisc_1(X0,X1),X2)) | (~r2_hidden(X1,X2) | ~r2_hidden(X0,X2))) | (~v1_classes2(X2) | v1_xboole_0(X2))) [ennf transformation 69]
121. ! [X0,X1,X2] : ((r2_hidden(k1_funct_2(X0,X1),X2) & r2_hidden(k2_zfmisc_1(X0,X1),X2)) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~v1_classes2(X2) | v1_xboole_0(X2)) [flattening 120]
137. ! [X0,X1] : (~v1_xboole_0(k2_zfmisc_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 31]
138. ! [X0,X1] : (~v1_xboole_0(k2_zfmisc_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 137]
146. ? [X0] : (? [X1] : ((? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X0,X1) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_funct_2(X2,X0,X0) & v1_funct_1(X2)) | ? [X3] : (~r2_hidden(X3,X1) & r2_hidden(X0,X1) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X0),X0))) & v1_funct_2(X3,k2_zfmisc_1(X0,X0),X0) & v1_funct_1(X3))) & v1_classes2(X1) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) => (? [X1] : ((? [X2] : (~r2_hidden(X2,X1) & r2_hidden(sK0,X1) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v1_funct_2(X2,sK0,sK0) & v1_funct_1(X2)) | ? [X3] : (~r2_hidden(X3,X1) & r2_hidden(sK0,X1) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) & v1_funct_2(X3,k2_zfmisc_1(sK0,sK0),sK0) & v1_funct_1(X3))) & v1_classes2(X1) & ~v1_xboole_0(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
147. ? [X1] : ((? [X2] : (~r2_hidden(X2,X1) & r2_hidden(sK0,X1) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v1_funct_2(X2,sK0,sK0) & v1_funct_1(X2)) | ? [X3] : (~r2_hidden(X3,X1) & r2_hidden(sK0,X1) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) & v1_funct_2(X3,k2_zfmisc_1(sK0,sK0),sK0) & v1_funct_1(X3))) & v1_classes2(X1) & ~v1_xboole_0(X1)) => ((? [X2] : (~r2_hidden(X2,sK1) & r2_hidden(sK0,sK1) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v1_funct_2(X2,sK0,sK0) & v1_funct_1(X2)) | ? [X3] : (~r2_hidden(X3,sK1) & r2_hidden(sK0,sK1) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) & v1_funct_2(X3,k2_zfmisc_1(sK0,sK0),sK0) & v1_funct_1(X3))) & v1_classes2(sK1) & ~v1_xboole_0(sK1)) [choice axiom]
148. ? [X2] : (~r2_hidden(X2,sK1) & r2_hidden(sK0,sK1) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v1_funct_2(X2,sK0,sK0) & v1_funct_1(X2)) => (~r2_hidden(sK2,sK1) & r2_hidden(sK0,sK1) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v1_funct_2(sK2,sK0,sK0) & v1_funct_1(sK2)) [choice axiom]
149. ? [X3] : (~r2_hidden(X3,sK1) & r2_hidden(sK0,sK1) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) & v1_funct_2(X3,k2_zfmisc_1(sK0,sK0),sK0) & v1_funct_1(X3)) => (~r2_hidden(sK3,sK1) & r2_hidden(sK0,sK1) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) & v1_funct_2(sK3,k2_zfmisc_1(sK0,sK0),sK0) & v1_funct_1(sK3)) [choice axiom]
150. (((~r2_hidden(sK2,sK1) & r2_hidden(sK0,sK1) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v1_funct_2(sK2,sK0,sK0) & v1_funct_1(sK2)) | (~r2_hidden(sK3,sK1) & r2_hidden(sK0,sK1) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) & v1_funct_2(sK3,k2_zfmisc_1(sK0,sK0),sK0) & v1_funct_1(sK3))) & v1_classes2(sK1) & ~v1_xboole_0(sK1)) & ~v1_xboole_0(sK0) [skolemisation 93,149,148,147,146]
175. ~v1_xboole_0(sK0) [cnf transformation 150] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
177. v1_classes2(sK1) [cnf transformation 150] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
178. v1_funct_1(sK2) | v1_funct_1(sK3) [cnf transformation 150] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
179. v1_funct_1(sK2) | v1_funct_2(sK3,k2_zfmisc_1(sK0,sK0),sK0) [cnf transformation 150] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
180. v1_funct_1(sK2) | m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) [cnf transformation 150] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
182. v1_funct_1(sK2) | ~r2_hidden(sK3,sK1) [cnf transformation 150] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
183. v1_funct_2(sK2,sK0,sK0) | v1_funct_1(sK3) [cnf transformation 150] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
184. v1_funct_2(sK2,sK0,sK0) | v1_funct_2(sK3,k2_zfmisc_1(sK0,sK0),sK0) [cnf transformation 150] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
185. v1_funct_2(sK2,sK0,sK0) | m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) [cnf transformation 150] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
187. v1_funct_2(sK2,sK0,sK0) | ~r2_hidden(sK3,sK1) [cnf transformation 150] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
188. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) | v1_funct_1(sK3) [cnf transformation 150] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
189. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) | v1_funct_2(sK3,k2_zfmisc_1(sK0,sK0),sK0) [cnf transformation 150] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
190. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) | m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) [cnf transformation 150] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
192. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) | ~r2_hidden(sK3,sK1) [cnf transformation 150] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
196. r2_hidden(sK0,sK1) | r2_hidden(sK0,sK1) [cnf transformation 150] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
198. ~r2_hidden(sK2,sK1) | v1_funct_1(sK3) [cnf transformation 150] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
199. ~r2_hidden(sK2,sK1) | v1_funct_2(sK3,k2_zfmisc_1(sK0,sK0),sK0) [cnf transformation 150] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
200. ~r2_hidden(sK2,sK1) | m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) [cnf transformation 150] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
202. ~r2_hidden(sK2,sK1) | ~r2_hidden(sK3,sK1) [cnf transformation 150] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
206. r2_hidden(X3,X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~v1_funct_2(X3,X0,X1) | ~v1_funct_1(X3) | ~v1_classes2(X2) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 95] {a:0,w:29,wCS:11600,sine:1,thAx:0,allAx:1,thDist:-1}
226. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 117] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
228. r2_hidden(k2_zfmisc_1(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~v1_classes2(X2) | v1_xboole_0(X2) [cnf transformation 121] {a:0,w:15,wCS:6000,sine:1,thAx:0,allAx:1,thDist:-1}
248. ~v1_xboole_0(k2_zfmisc_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 138] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
266. r2_hidden(sK0,sK1) [duplicate literal removal 196] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
268. 1 <=> r2_hidden(sK3,sK1) [avatar definition]
270. ~r2_hidden(sK3,sK1) <- (~1) [avatar component clause 268] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
272. 2 <=> r2_hidden(sK2,sK1) [avatar definition]
274. ~r2_hidden(sK2,sK1) <- (~2) [avatar component clause 272] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
275. ~1 | ~2 [avatar split clause 202,272,268]
277. 3 <=> r2_hidden(sK0,sK1) [avatar definition]
279. r2_hidden(sK0,sK1) <- (3) [avatar component clause 277] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
282. 4 <=> m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) [avatar definition]
284. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK0),sK0))) <- (4) [avatar component clause 282] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
285. 4 | ~2 [avatar split clause 200,272,282]
287. 5 <=> v1_funct_2(sK3,k2_zfmisc_1(sK0,sK0),sK0) [avatar definition]
289. v1_funct_2(sK3,k2_zfmisc_1(sK0,sK0),sK0) <- (5) [avatar component clause 287] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
290. 5 | ~2 [avatar split clause 199,272,287]
292. 6 <=> v1_funct_1(sK3) [avatar definition]
294. v1_funct_1(sK3) <- (6) [avatar component clause 292] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
295. 6 | ~2 [avatar split clause 198,272,292]
297. 3 [avatar split clause 266,277]
302. 7 <=> m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) [avatar definition]
304. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) <- (7) [avatar component clause 302] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
305. ~1 | 7 [avatar split clause 192,302,268]
307. 4 | 7 [avatar split clause 190,302,282]
308. 5 | 7 [avatar split clause 189,302,287]
309. 6 | 7 [avatar split clause 188,302,292]
311. 8 <=> v1_funct_2(sK2,sK0,sK0) [avatar definition]
313. v1_funct_2(sK2,sK0,sK0) <- (8) [avatar component clause 311] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
314. ~1 | 8 [avatar split clause 187,311,268]
316. 4 | 8 [avatar split clause 185,311,282]
317. 5 | 8 [avatar split clause 184,311,287]
318. 6 | 8 [avatar split clause 183,311,292]
320. 9 <=> v1_funct_1(sK2) [avatar definition]
322. v1_funct_1(sK2) <- (9) [avatar component clause 320] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
323. ~1 | 9 [avatar split clause 182,320,268]
325. 4 | 9 [avatar split clause 180,320,282]
326. 5 | 9 [avatar split clause 179,320,287]
327. 6 | 9 [avatar split clause 178,320,292]
329. 10 <=> v1_classes2(sK1) [avatar definition]
331. v1_classes2(sK1) <- (10) [avatar component clause 329] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
332. 10 [avatar split clause 177,329]
339. 12 <=> v1_xboole_0(sK0) [avatar definition]
341. ~v1_xboole_0(sK0) <- (~12) [avatar component clause 339] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
342. ~12 [avatar split clause 175,339]
534. 28 <=> v1_xboole_0(k2_zfmisc_1(sK0,sK0)) [avatar definition]
535. v1_xboole_0(k2_zfmisc_1(sK0,sK0)) <- (28) [avatar component clause 534] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
596. r2_hidden(k2_zfmisc_1(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~v1_classes2(X2) [subsumption resolution 228,226] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
786. ~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | r2_hidden(X3,X2) | ~v1_funct_2(X3,X0,X1) | ~v1_funct_1(X3) | ~v1_classes2(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [subsumption resolution 206,226] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
787. $false <- (~2, 3, 7, 8, 9, 10, ~12) [unit resulting resolution 341,341,331,322,279,274,279,313,304,786] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:46,thDist:-46}
790. ~r2_hidden(sK0,X5) | ~r2_hidden(k2_zfmisc_1(sK0,sK0),X5) | r2_hidden(sK3,X5) | ~v1_funct_2(sK3,k2_zfmisc_1(sK0,sK0),sK0) | ~v1_funct_1(sK3) | ~v1_classes2(X5) | v1_xboole_0(sK0) | v1_xboole_0(k2_zfmisc_1(sK0,sK0)) <- (4) [resolution 786,284] {a:1,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
800. 2 | ~3 | ~7 | ~8 | ~9 | ~10 | 12 [avatar contradiction clause 787]
810. ~r2_hidden(sK0,X5) | ~r2_hidden(k2_zfmisc_1(sK0,sK0),X5) | r2_hidden(sK3,X5) | ~v1_funct_1(sK3) | ~v1_classes2(X5) | v1_xboole_0(sK0) | v1_xboole_0(k2_zfmisc_1(sK0,sK0)) <- (4, 5) [subsumption resolution 790,289] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
811. ~r2_hidden(sK0,X5) | ~r2_hidden(k2_zfmisc_1(sK0,sK0),X5) | r2_hidden(sK3,X5) | ~v1_funct_1(sK3) | ~v1_classes2(X5) | v1_xboole_0(k2_zfmisc_1(sK0,sK0)) <- (4, 5, ~12) [subsumption resolution 810,341] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
815. 47 <=> ! [X5] : (~r2_hidden(sK0,X5) | ~v1_classes2(X5) | r2_hidden(sK3,X5)) [avatar definition]
816. ~r2_hidden(sK0,X5) | ~v1_classes2(X5) | r2_hidden(sK3,X5) <- (47) [avatar component clause 815] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
922. ~v1_classes2(sK1) | r2_hidden(sK3,sK1) <- (3, 47) [resolution 816,279] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
925. r2_hidden(sK3,sK1) <- (3, 10, 47) [subsumption resolution 922,331] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
926. $false <- (~1, 3, 10, 47) [subsumption resolution 925,270] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
927. 1 | ~3 | ~10 | ~47 [avatar contradiction clause 926]
946. ~r2_hidden(sK0,X5) | ~r2_hidden(k2_zfmisc_1(sK0,sK0),X5) | r2_hidden(sK3,X5) | ~v1_classes2(X5) | v1_xboole_0(k2_zfmisc_1(sK0,sK0)) <- (4, 5, 6, ~12) [subsumption resolution 811,294] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
947. ~r2_hidden(sK0,X5) | r2_hidden(sK3,X5) | ~v1_classes2(X5) | v1_xboole_0(k2_zfmisc_1(sK0,sK0)) <- (4, 5, 6, ~12) [subsumption resolution 946,596] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
948. 28 | 47 | ~4 | ~5 | ~6 | 12 [avatar split clause 947,339,292,287,282,815,534]
963. v1_xboole_0(sK0) | v1_xboole_0(sK0) <- (28) [resolution 535,248] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
964. v1_xboole_0(sK0) <- (28) [duplicate literal removal 963] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
966. $false <- (~12, 28) [subsumption resolution 964,341] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
967. 12 | ~28 [avatar contradiction clause 966]
968. $false [avatar sat refutation 275,285,290,295,297,305,307,308,309,314,316,317,318,323,325,326,327,332,342,800,927,948,967] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:861,thDist:-861}
% SZS output end Proof for mod_2__l57_mod_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 59
% >>> Preprocessing
% Introduced skolems: 16
% Pure predicates: 5
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 90
% Generated clauses: 540
% Active clauses: 116
% Passive clauses: 272
% Final active clauses: 101
% Final passive clauses: 149
% Main loop iterations started: 121
% 
% >>> Simplifying Inferences
% Duplicate literals: 10
% Fw subsumption resolutions: 91
% 
% >>> Deletion Inferences
% Simple tautologies: 12
% Forward subsumptions: 56
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 238
% Unit resulting resolution: 3
% Forward superposition: 5
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 55
% Split components: 114
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 97
% SAT solver unit clauses: 9
% SAT solver binary clauses: 41
% 
% Memory used [KB]: 9978
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s ( own 0.0000 s ) 
% sine selection: 0.001 s
% forward subsumption resolution: 0.001 s
% resolution: 0.001 s
% unit resulting resolution: 0.002 s
% superposition: 0.001 s
% other: 0.007 s ( own 0.001 s ) 

