i: [1,38,0,2,1,1,0,0] t102_relat_1
a: [3,70,0,2,1,38]
s: 70
i: [1,37,0,2,1,1,0,0] t102_relat_1
a: [3,75,0,2,1,37]
s: 75
i: [1,40,0,3,1,1,0,0] t102_relat_1
a: [3,60,0,3,1,40]
s: 60
i: [1,39,0,3,1,1,0,0] t102_relat_1
a: [3,65,0,3,1,39]
s: 65
i: [1,41,0,7,1,1,0,0] t102_relat_1
a: [3,55,0,7,1,41]
s: 55
i: [1,49,0,3,1,0,0,1] reflexivity_r1_tarski
s: 49
i: [1,48,0,4,1,0,0,255] existence_m1_subset_1
s: 48
i: [1,50,0,6,2,0,0,1] dt_k6_relat_1
s: 50
i: [1,44,0,7,2,0,0,255] t3_subset
s: 44
i: [1,45,0,7,2,0,0,255] t3_subset
s: 45
i: [1,42,0,8,3,0,0,255] cc2_relat_1
s: 42
d: [2,77,1,5,1,0,40,44,48]
s: 77
i: [1,47,0,9,3,0,0,1] t1_xboole_1
s: 47
d: [2,82,1,6,2,1,40,47,60]
s: 82
i: [1,51,0,11,2,0,0,255] rd7_relat_1
s: 51
i: [1,46,0,12,3,0,0,1] t100_relat_1
s: 46
d: [2,86,2,5,1,1,40,82,77]
a: [3,90,2,5,1,86]
s: 90
i: [1,43,0,14,4,0,0,1] t101_relat_1
s: 43
d: [2,83,1,6,2,1,40,47,65]
s: 83
d: [2,79,1,6,2,0,40,42,48]
s: 79
d: [2,80,1,7,3,0,40,42,45]
s: 80
d: [2,115,2,4,2,1,40,80,60]
a: [3,126,2,2,1,115]
s: 126
d: [2,95,1,18,4,0,44,46,51]
d: [2,97,1,14,3,0,21,95,50]
s: 97
d: [2,106,2,5,1,1,40,83,77]
a: [3,110,2,5,1,106]
s: 110
d: [2,96,1,18,4,0,44,46,51]
d: [2,98,1,14,3,0,21,96,50]
s: 98
d: [2,94,1,15,4,0,40,46,47]
s: 94
d: [2,138,3,6,2,1,40,110,80]
d: [2,140,3,4,1,2,21,138,70]
a: [3,144,3,4,1,140]
s: 144
d: [2,101,1,20,5,0,44,43,51]
d: [2,103,1,16,4,0,21,101,50]
s: 103
d: [2,99,3,8,2,1,40,90,47]
s: 99
d: [2,102,1,20,5,0,44,43,51]
d: [2,104,1,16,4,0,21,102,50]
s: 104
d: [2,100,1,17,5,0,40,43,47]
s: 100
d: [2,171,4,7,1,1,40,99,77]
a: [3,175,4,7,1,171]
s: 175
d: [2,183,1,15,4,0,21,100,80]
s: 183
d: [2,139,3,8,2,1,40,110,47]
s: 139
d: [2,84,2,8,2,0,40,47,77]
s: 84
d: [2,186,2,12,3,0,40,183,49]
s: 186
r: 43
d: [2,201,4,7,1,1,40,139,77]
a: [3,205,4,7,1,201]
s: 205
d: [2,156,2,14,3,0,40,94,77]
s: 156
d: [2,219,5,8,2,1,40,205,80]
d: [2,221,5,6,1,2,21,219,70]
a: [3,225,5,6,1,221]
s: 225
d: [2,193,2,14,3,0,40,183,77]
s: 193
d: [2,152,2,17,5,0,40,94,46]
d: [2,160,2,15,4,0,18,152]
s: 160
d: [2,185,5,10,2,1,40,175,47]
s: 185
d: [2,154,2,19,6,0,40,94,43]
d: [2,159,2,17,5,0,18,154]
d: [2,162,2,15,4,0,21,159,80]
s: 162
d: [2,265,3,0,0,4,52,70,65,60,55,162]
e: 265
d: [2,266,3,8,3,1,40,162,55]
d: [2,275,3,6,2,2,21,266,70]
d: [2,276,3,3,1,3,21,275,60]
d: [2,277,3,0,0,4,21,276,65]
e: 277
f: 277,38,39,40,41
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t102_relat_1
% SZS output start Proof for relat_1__t102_relat_1
1. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => ! [X3] : (v1_relat_1(X3) => ((r1_tarski(X0,X1) & r1_tarski(X2,X3)) => r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X3))))) [input t102_relat_1]
2. ~! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => ! [X3] : (v1_relat_1(X3) => ((r1_tarski(X0,X1) & r1_tarski(X2,X3)) => r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X3))))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
10. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X2)))) [input t100_relat_1]
11. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2))))) [input t101_relat_1]
12. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
13. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
14. ~! [X0,X1,X2] : (v1_relat_1(X2) => ! [X3] : (v1_relat_1(X3) => ((r1_tarski(X0,X1) & r1_tarski(X2,X3)) => r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X3))))) [flattening 2]
15. ! [X0,X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2))))) [flattening 11]
16. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 13]
17. ! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X2)))) [flattening 10]
18. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 12]
20. ? [X0,X1,X2] : (? [X3] : ((~r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X3)) & (r1_tarski(X0,X1) & r1_tarski(X2,X3))) & v1_relat_1(X3)) & v1_relat_1(X2)) [ennf transformation 14]
21. ? [X0,X1,X2] : (? [X3] : (~r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X3)) & r1_tarski(X0,X1) & r1_tarski(X2,X3) & v1_relat_1(X3)) & v1_relat_1(X2)) [flattening 20]
22. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
23. ! [X0,X1] : (! [X2] : ((r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2)) | ~r1_tarski(X1,X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [ennf transformation 15]
24. ! [X0,X1] : (! [X2] : (r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [flattening 23]
25. ! [X0,X1,X2] : ((r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X2)) | ~r1_tarski(X0,X1)) | ~v1_relat_1(X2)) [ennf transformation 17]
26. ! [X0,X1,X2] : (r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X2)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X2)) [flattening 25]
27. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 18]
28. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 27]
31. ? [X0,X1,X2] : (? [X3] : (~r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X3)) & r1_tarski(X0,X1) & r1_tarski(X2,X3) & v1_relat_1(X3)) & v1_relat_1(X2)) => (? [X3] : (~r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK1,X3)) & r1_tarski(sK0,sK1) & r1_tarski(sK2,X3) & v1_relat_1(X3)) & v1_relat_1(sK2)) [choice axiom]
32. ? [X3] : (~r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK1,X3)) & r1_tarski(sK0,sK1) & r1_tarski(sK2,X3) & v1_relat_1(X3)) => (~r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK1,sK3)) & r1_tarski(sK0,sK1) & r1_tarski(sK2,sK3) & v1_relat_1(sK3)) [choice axiom]
33. (~r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK1,sK3)) & r1_tarski(sK0,sK1) & r1_tarski(sK2,sK3) & v1_relat_1(sK3)) & v1_relat_1(sK2) [skolemisation 21,32,31]
34. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 16]
38. v1_relat_1(sK3) [cnf transformation 33] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. r1_tarski(sK2,sK3) [cnf transformation 33] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. r1_tarski(sK0,sK1) [cnf transformation 33] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. ~r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK1,sK3)) [cnf transformation 33] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 24] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 34] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
46. r1_tarski(k6_relat_1(X0,X2),k6_relat_1(X1,X2)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X2) [cnf transformation 26] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. 1 <=> r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK1,sK3)) [avatar definition]
55. ~r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK1,sK3)) <- (~1) [avatar component clause 53] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
56. ~1 [avatar split clause 41,53]
58. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
60. r1_tarski(sK0,sK1) <- (2) [avatar component clause 58] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
61. 2 [avatar split clause 40,58]
63. 3 <=> r1_tarski(sK2,sK3) [avatar definition]
65. r1_tarski(sK2,sK3) <- (3) [avatar component clause 63] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66. 3 [avatar split clause 39,63]
68. 4 <=> v1_relat_1(sK3) [avatar definition]
70. v1_relat_1(sK3) <- (4) [avatar component clause 68] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
71. 4 [avatar split clause 38,68]
80. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 42,45] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
94. ~r1_tarski(X3,k6_relat_1(X0,X2)) | ~v1_relat_1(X2) | r1_tarski(X3,k6_relat_1(X1,X2)) | ~r1_tarski(X0,X1) [resolution 46,47] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
154. ~v1_relat_1(X11) | r1_tarski(k6_relat_1(X12,X13),k6_relat_1(X14,X11)) | ~r1_tarski(X12,X14) | ~r1_tarski(X13,X11) | ~v1_relat_1(X11) | ~v1_relat_1(X13) [resolution 94,43] {a:2,w:19,wCS:7600,sine:1,thAx:0,allAx:3,thDist:-3}
159. ~v1_relat_1(X11) | r1_tarski(k6_relat_1(X12,X13),k6_relat_1(X14,X11)) | ~r1_tarski(X12,X14) | ~r1_tarski(X13,X11) | ~v1_relat_1(X13) [duplicate literal removal 154] {a:2,w:17,wCS:6800,sine:1,thAx:0,allAx:3,thDist:-3}
162. r1_tarski(k6_relat_1(X12,X13),k6_relat_1(X14,X11)) | ~v1_relat_1(X11) | ~r1_tarski(X12,X14) | ~r1_tarski(X13,X11) [subsumption resolution 159,80] {a:2,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
266. ~v1_relat_1(sK3) | ~r1_tarski(sK0,sK1) | ~r1_tarski(sK2,sK3) <- (~1) [resolution 162,55] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
275. ~r1_tarski(sK0,sK1) | ~r1_tarski(sK2,sK3) <- (~1, 4) [subsumption resolution 266,70] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
276. ~r1_tarski(sK2,sK3) <- (~1, 2, 4) [subsumption resolution 275,60] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
277. $false <- (~1, 2, 3, 4) [subsumption resolution 276,65] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
278. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 277]
282. $false [avatar sat refutation 56,61,66,71,278] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:150,thDist:-150}
% SZS output end Proof for relat_1__t102_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 182
% Active clauses: 42
% Passive clauses: 83
% Final active clauses: 41
% Final passive clauses: 40
% Main loop iterations started: 43
% 
% >>> Simplifying Inferences
% Duplicate literals: 12
% Fw subsumption resolutions: 34
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Forward subsumptions: 26
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 81
% Unit resulting resolution: 1
% Forward superposition: 24
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 5
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

