i: [1,128,0,4,1,1,0,0] l8_arithm
a: [3,199,0,4,1,128]
s: 199
i: [1,129,0,2,1,0,0,1] spc1_boole
a: [3,204,0,2,1,129]
s: 204
i: [1,136,0,2,1,0,0,255] fc6_ordinal1
a: [3,234,0,2,1,136]
s: 234
i: [1,135,0,2,1,0,0,255] fc6_ordinal1
a: [3,239,0,2,1,135]
s: 239
i: [1,191,0,8,3,0,0,255] d7_xcmplx_0
d: [2,195,0,6,2,0,18,191]
a: [3,259,0,2,1,195]
s: 259
i: [1,179,0,2,1,0,0,255] rc4_ordinal1
a: [3,273,0,2,1,179]
s: 273
i: [1,180,0,2,1,0,0,255] rc1_ordinal1
a: [3,278,0,2,1,180]
s: 278
i: [1,181,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,283,0,2,1,181]
s: 283
i: [1,185,0,2,1,0,0,255] rc2_ordinal1
a: [3,288,0,2,1,185]
s: 288
i: [1,184,0,2,1,0,0,255] rc2_ordinal1
a: [3,293,0,2,1,184]
s: 293
i: [1,183,0,2,1,0,0,255] rc2_ordinal1
a: [3,298,0,2,1,183]
s: 298
i: [1,182,0,2,1,0,0,255] rc2_ordinal1
a: [3,303,0,2,1,182]
s: 303
i: [1,187,0,2,1,0,0,255] rc5_ordinal1
a: [3,308,0,2,1,187]
s: 308
i: [1,186,0,2,1,0,0,255] rc5_ordinal1
a: [3,313,0,2,1,186]
s: 313
i: [1,189,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,318,0,2,1,189]
s: 318
i: [1,188,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,323,0,2,1,188]
s: 323
i: [1,170,0,3,1,0,0,255] reflexivity_r1_tarski
s: 170
i: [1,130,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,209,0,3,1,130]
s: 209
i: [1,131,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,214,0,3,1,131]
s: 214
i: [1,132,0,3,1,0,0,255] t1_numerals
a: [3,219,0,3,1,132]
s: 219
i: [1,139,0,3,1,0,0,255] spc1_numerals
a: [3,244,0,3,1,139]
s: 244
i: [1,138,0,3,1,0,0,255] spc1_numerals
a: [3,249,0,3,1,138]
s: 249
d: [2,325,0,3,1,2,22,219,209]
a: [3,329,0,3,1,325]
s: 329
d: [2,331,0,3,1,2,22,249,214]
a: [3,335,0,3,1,331]
s: 335
i: [1,141,0,4,2,0,0,255] cc6_ordinal1
s: 141
i: [1,142,0,4,2,0,0,255] cc2_xcmplx_0
s: 142
d: [2,337,1,2,1,1,40,141,273]
a: [3,342,1,2,1,337]
s: 342
i: [1,143,0,4,2,0,0,255] cc1_ordinal1
s: 143
d: [2,338,1,2,1,1,40,141,308]
a: [3,347,1,2,1,338]
s: 347
i: [1,144,0,4,2,0,0,255] cc1_ordinal1
s: 144
i: [1,153,0,4,2,0,0,255] cc9_ordinal1
s: 153
d: [2,349,1,2,1,1,40,142,273]
a: [3,354,1,2,1,349]
s: 354
i: [1,154,0,4,2,0,0,255] cc7_ordinal1
s: 154
d: [2,350,1,2,1,1,40,142,308]
a: [3,359,1,2,1,350]
s: 359
i: [1,155,0,4,2,0,0,255] cc3_ordinal1
s: 155
i: [1,169,0,4,1,0,0,255] existence_m1_subset_1
s: 169
d: [2,361,1,2,1,1,40,143,234]
a: [3,368,1,2,1,361]
s: 368
i: [1,133,0,4,1,0,0,255] dt_k5_numbers
a: [3,224,0,4,1,133]
s: 224
d: [2,363,1,2,1,1,40,143,278]
a: [3,378,1,2,1,363]
s: 378
i: [1,134,0,4,1,0,0,255] dt_k6_numbers
a: [3,229,0,4,1,134]
s: 229
i: [1,137,0,4,1,0,0,255] spc1_numerals
a: [3,254,0,4,1,137]
s: 254
d: [2,386,1,2,1,1,40,144,234]
a: [3,394,1,2,1,386]
s: 394
d: [2,411,0,4,1,2,22,224,214]
a: [3,415,0,4,1,411]
s: 415
d: [2,388,1,2,1,1,40,144,278]
a: [3,404,1,2,1,388]
s: 404
d: [2,417,0,4,1,2,22,229,214]
a: [3,421,0,4,1,417]
s: 421
d: [2,423,0,4,1,2,22,254,214]
a: [3,427,0,4,1,423]
s: 427
d: [2,362,2,2,1,1,40,143,342]
a: [3,373,2,2,1,362]
s: 373
i: [1,146,0,5,2,0,0,1] dt_k5_xcmplx_0
s: 146
d: [2,380,2,2,1,1,40,347,143]
a: [3,384,2,2,1,380]
s: 384
i: [1,156,0,5,2,0,0,255] t6_boole
s: 156
i: [1,157,0,5,2,0,0,255] cc8_ordinal1
s: 157
d: [2,430,1,2,1,1,40,157,335]
a: [3,436,1,2,1,430]
s: 436
i: [1,158,0,5,2,0,0,255] cc1_xcmplx_0
s: 158
d: [2,431,1,2,1,1,40,157,329]
a: [3,441,1,2,1,431]
s: 441
d: [2,468,2,0,0,2,52,259,441,142]
e: 468
d: [2,469,2,2,1,1,40,441,142]
d: [2,472,2,0,0,2,21,469,259]
e: 472
r: 259
a: [3,258,0,2,1,195]
s: 258
a: [3,263,0,4,1,195]
s: 263
d: [2,387,2,2,1,1,40,144,342]
a: [3,399,2,2,1,387]
s: 399
i: [1,166,0,5,2,0,0,255] t7_boole
s: 166
d: [2,390,2,2,1,1,40,144,347]
a: [3,409,2,2,1,390]
s: 409
d: [2,429,0,5,2,1,22,156,209]
s: 429
i: [1,161,0,6,3,0,0,255] cc2_ordinal1
s: 161
d: [2,448,2,2,1,1,40,436,142]
a: [3,453,2,2,1,448]
s: 453
i: [1,162,0,6,2,0,0,255] t1_subset
s: 162
d: [2,449,2,2,1,1,40,436,141]
a: [3,458,2,2,1,449]
s: 458
i: [1,171,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 171
i: [1,145,0,7,3,0,0,255] cc5_ordinal1
s: 145
d: [2,470,2,2,1,1,40,441,141]
a: [3,478,2,2,1,470]
s: 478
i: [1,147,0,7,2,0,0,1] involutiveness_k5_xcmplx_0
s: 147
d: [2,492,3,2,1,1,40,458,144]
a: [3,497,3,2,1,492]
s: 497
i: [1,148,0,7,2,0,0,1] t3_arithm
s: 148
i: [1,159,0,7,3,0,0,1] fc7_xcmplx_0
s: 159
d: [2,493,3,2,1,1,40,458,143]
a: [3,502,3,2,1,493]
s: 502
i: [1,160,0,7,3,0,0,1] fc7_xcmplx_0
s: 160
d: [2,511,3,2,1,1,40,478,144]
a: [3,517,3,2,1,511]
s: 517
i: [1,164,0,7,2,0,0,255] t3_subset
s: 164
i: [1,165,0,7,3,0,0,1] t8_boole
s: 165
d: [2,512,3,2,1,1,40,478,143]
a: [3,522,3,2,1,512]
s: 522
i: [1,140,0,8,3,0,0,255] cc10_ordinal1
s: 140
d: [2,535,1,4,2,1,40,140,415]
a: [3,540,1,2,1,535]
s: 540
i: [1,163,0,8,3,0,0,255] t2_subset
s: 163
i: [1,177,0,8,3,0,0,2] fc3_xcmplx_0
s: 177
d: [2,432,1,3,1,0,40,157,169]
a: [3,446,1,3,1,432]
s: 446
i: [1,168,0,9,3,0,0,255] t5_subset
s: 168
d: [2,547,1,5,2,1,40,163,244]
a: [3,558,1,2,1,547]
s: 558
i: [1,167,0,10,3,0,0,255] t4_subset
s: 167
i: [1,178,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 178
d: [2,461,1,3,1,0,40,158,169]
a: [3,466,1,3,1,461]
s: 466
i: [1,176,0,12,5,0,0,2] fc8_xcmplx_0
s: 176
a: [3,555,1,3,1,547]
s: 555
i: [1,175,0,14,4,0,0,255] existence_m2_subset_1
s: 175
i: [1,194,0,14,4,0,0,255] d7_xcmplx_0
s: 194
d: [2,548,1,5,2,1,40,163,335]
d: [2,561,1,3,1,2,21,548,239]
a: [3,565,1,3,1,561]
s: 565
d: [2,636,0,11,3,0,21,194,146]
s: 636
d: [2,647,1,12,4,0,44,636,148]
d: [2,654,1,8,3,1,21,647,199]
d: [2,655,1,6,2,2,21,654,453]
a: [3,659,1,3,1,655]
s: 659
d: [2,668,3,0,0,2,52,453,659,146]
e: 668
d: [2,669,2,2,1,1,40,659,146]
d: [2,671,2,0,0,2,21,669,453]
e: 671
d: [2,460,1,2,1,1,40,158,244]
d: [2,648,1,12,4,0,44,148,636]
d: [2,665,1,8,3,1,21,648,199]
d: [2,666,1,6,2,2,21,665,453]
f: 671,666,469,460,431,325,195,139,132,130,128
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arithm__l8_arithm
% SZS output start Proof for arithm__l8_arithm
1. 1 = k5_xcmplx_0(1) [input l8_arithm]
2. ~1 = k5_xcmplx_0(1) [negated conjecture 1]
6. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xcmplx_0(X0)) [input cc1_xcmplx_0]
8. ! [X0] : (v7_ordinal1(X0) => v1_xcmplx_0(X0)) [input cc2_xcmplx_0]
14. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
17. ! [X0] : (v1_xcmplx_0(X0) => ! [X1] : (v1_xcmplx_0(X1) => ((k6_numbers = X0 => (k5_xcmplx_0(X0) = X1 <=> k6_numbers = X1)) & (k6_numbers != X0 => (k5_xcmplx_0(X0) = X1 <=> 1 = k3_xcmplx_0(X0,X1)))))) [input d7_xcmplx_0]
24. ! [X0] : (v1_xcmplx_0(X0) => v1_xcmplx_0(k5_xcmplx_0(X0))) [input dt_k5_xcmplx_0]
42. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
46. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
47. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
50. ! [X0] : (v1_xcmplx_0(X0) => k3_xcmplx_0(1,X0) = X0) [input t3_arithm]
57. 1 != k5_xcmplx_0(1) [flattening 2]
67. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) [pure predicate removal 46]
71. ! [X0] : (v1_xcmplx_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 8]
74. ! [X0] : (v1_xcmplx_0(k5_xcmplx_0(X0)) | ~v1_xcmplx_0(X0)) [ennf transformation 24]
76. ! [X0] : (k3_xcmplx_0(1,X0) = X0 | ~v1_xcmplx_0(X0)) [ennf transformation 50]
77. ! [X0] : (! [X1] : ((((k5_xcmplx_0(X0) = X1 <=> k6_numbers = X1) | k6_numbers != X0) & ((k5_xcmplx_0(X0) = X1 <=> 1 = k3_xcmplx_0(X0,X1)) | k6_numbers = X0)) | ~v1_xcmplx_0(X1)) | ~v1_xcmplx_0(X0)) [ennf transformation 17]
82. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 14]
83. ! [X0] : (v1_xcmplx_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 6]
110. ! [X0] : (! [X1] : (((((k5_xcmplx_0(X0) = X1 | k6_numbers != X1) & (k6_numbers = X1 | k5_xcmplx_0(X0) != X1)) | k6_numbers != X0) & (((k5_xcmplx_0(X0) = X1 | 1 != k3_xcmplx_0(X0,X1)) & (1 = k3_xcmplx_0(X0,X1) | k5_xcmplx_0(X0) != X1)) | k6_numbers = X0)) | ~v1_xcmplx_0(X1)) | ~v1_xcmplx_0(X0)) [nnf transformation 77]
128. 1 != k5_xcmplx_0(1) [cnf transformation 57] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
130. k6_numbers = k1_xboole_0 [cnf transformation 42] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
132. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 47] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
139. m1_subset_1(1,k1_numbers) [cnf transformation 67] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
142. ~v7_ordinal1(X0) | v1_xcmplx_0(X0) [cnf transformation 71] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
146. v1_xcmplx_0(k5_xcmplx_0(X0)) | ~v1_xcmplx_0(X0) [cnf transformation 74] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
148. k3_xcmplx_0(1,X0) = X0 | ~v1_xcmplx_0(X0) [cnf transformation 76] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
149. 1 = k3_xcmplx_0(X0,X1) | k5_xcmplx_0(X0) != X1 | k6_numbers = X0 | ~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0) [cnf transformation 110] {a:0,w:16,wCS:6400,sine:255,thAx:0,allAx:1,thDist:-1}
152. k5_xcmplx_0(X0) = X1 | k6_numbers != X1 | k6_numbers != X0 | ~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0) [cnf transformation 110] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
157. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 82] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
158. ~m1_subset_1(X0,k1_numbers) | v1_xcmplx_0(X0) [cnf transformation 83] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
190. k6_numbers = k5_xcmplx_0(X0) | k6_numbers != X0 | ~v1_xcmplx_0(k6_numbers) | ~v1_xcmplx_0(X0) [equality resolution 152] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
191. k6_numbers = k5_xcmplx_0(k6_numbers) | ~v1_xcmplx_0(k6_numbers) | ~v1_xcmplx_0(k6_numbers) [equality resolution 190] {a:0,w:8,wCS:3200,sine:255,thAx:0,allAx:1,thDist:-1}
194. 1 = k3_xcmplx_0(X0,k5_xcmplx_0(X0)) | k6_numbers = X0 | ~v1_xcmplx_0(k5_xcmplx_0(X0)) | ~v1_xcmplx_0(X0) [equality resolution 149] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
195. k6_numbers = k5_xcmplx_0(k6_numbers) | ~v1_xcmplx_0(k6_numbers) [duplicate literal removal 191] {a:0,w:6,wCS:2400,sine:255,thAx:0,allAx:1,thDist:-1}
197. 1 <=> 1 = k5_xcmplx_0(1) [avatar definition]
199. 1 != k5_xcmplx_0(1) <- (~1) [avatar component clause 197] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
200. ~1 [avatar split clause 128,197]
207. 3 <=> k6_numbers = k1_xboole_0 [avatar definition]
209. k6_numbers = k1_xboole_0 <- (3) [avatar component clause 207] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
210. 3 [avatar split clause 130,207]
217. 5 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
219. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (5) [avatar component clause 217] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
220. 5 [avatar split clause 132,217]
242. 10 <=> m1_subset_1(1,k1_numbers) [avatar definition]
244. m1_subset_1(1,k1_numbers) <- (10) [avatar component clause 242] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
245. 10 [avatar split clause 139,242]
257. 13 <=> v1_xcmplx_0(k6_numbers) [avatar definition]
261. 14 <=> k6_numbers = k5_xcmplx_0(k6_numbers) [avatar definition]
264. ~13 | 14 [avatar split clause 195,261,257]
325. m1_subset_1(k6_numbers,k4_ordinal1) <- (3, 5) [forward demodulation 219,209] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
327. 27 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
329. m1_subset_1(k6_numbers,k4_ordinal1) <- (27) [avatar component clause 327] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
330. 27 | ~3 | ~5 [avatar split clause 325,217,207,327]
431. v7_ordinal1(k6_numbers) <- (27) [resolution 157,329] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
439. 45 <=> v7_ordinal1(k6_numbers) [avatar definition]
441. v7_ordinal1(k6_numbers) <- (45) [avatar component clause 439] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
442. 45 | ~27 [avatar split clause 431,327,439]
451. 47 <=> v1_xcmplx_0(1) [avatar definition]
453. v1_xcmplx_0(1) <- (47) [avatar component clause 451] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
460. v1_xcmplx_0(1) <- (10) [resolution 158,244] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
462. 47 | ~10 [avatar split clause 460,242,451]
469. v1_xcmplx_0(k6_numbers) <- (45) [resolution 441,142] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
474. 13 | ~45 [avatar split clause 469,439,257]
636. 1 = k3_xcmplx_0(X0,k5_xcmplx_0(X0)) | k6_numbers = X0 | ~v1_xcmplx_0(X0) [subsumption resolution 194,146] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
648. 1 = k5_xcmplx_0(1) | ~v1_xcmplx_0(k5_xcmplx_0(1)) | 1 = k6_numbers | ~v1_xcmplx_0(1) [superposition 148,636] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
657. 68 <=> v1_xcmplx_0(k5_xcmplx_0(1)) [avatar definition]
659. ~v1_xcmplx_0(k5_xcmplx_0(1)) <- (~68) [avatar component clause 657] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
661. 69 <=> 1 = k6_numbers [avatar definition]
665. ~v1_xcmplx_0(k5_xcmplx_0(1)) | 1 = k6_numbers | ~v1_xcmplx_0(1) <- (~1) [subsumption resolution 648,199] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
666. ~v1_xcmplx_0(k5_xcmplx_0(1)) | 1 = k6_numbers <- (~1, 47) [subsumption resolution 665,453] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
667. 69 | ~68 | 1 | ~47 [avatar split clause 666,451,197,657,661]
669. ~v1_xcmplx_0(1) <- (~68) [resolution 659,146] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
671. $false <- (47, ~68) [subsumption resolution 669,453] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
672. ~47 | 68 [avatar contradiction clause 671]
673. k6_numbers != k5_xcmplx_0(k6_numbers) | 1 != k6_numbers | 1 = k5_xcmplx_0(1) [theory tautology sat conflict] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
674. $false [avatar sat refutation 200,210,220,245,264,330,442,462,474,667,672,673] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:296,thDist:-296}
% SZS output end Proof for arithm__l8_arithm
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 55
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 62
% Generated clauses: 259
% Active clauses: 86
% Passive clauses: 128
% Final active clauses: 85
% Final passive clauses: 34
% Main loop iterations started: 94
% 
% >>> Simplifying Inferences
% Duplicate literals: 17
% Fw subsumption resolutions: 16
% Fw demodulations: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 26
% 
% >>> Generating Inferences
% Binary resolution: 64
% Unit resulting resolution: 2
% Forward superposition: 14
% Backward superposition: 14
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 17
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 75
% SAT solver unit clauses: 25
% SAT solver binary clauses: 32
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.003 s

