i: [1,39,0,2,1,1,0,0] t33_projpl_1
a: [3,90,0,2,1,39]
s: 90
i: [1,38,0,2,1,1,0,0] t33_projpl_1
a: [3,95,0,2,1,38]
s: 95
i: [1,37,0,2,1,1,0,0] t33_projpl_1
a: [3,100,0,2,1,37]
s: 100
i: [1,36,0,2,1,1,0,0] t33_projpl_1
a: [3,105,0,2,1,36]
s: 105
i: [1,35,0,2,1,1,0,0] t33_projpl_1
a: [3,110,0,2,1,35]
s: 110
i: [1,34,0,2,1,1,0,0] t33_projpl_1
a: [3,115,0,2,1,34]
s: 115
i: [1,44,0,3,1,1,0,0] t33_projpl_1
a: [3,65,0,3,1,44]
s: 65
i: [1,42,0,4,1,1,0,0] t33_projpl_1
a: [3,75,0,4,1,42]
s: 75
i: [1,41,0,4,1,1,0,0] t33_projpl_1
a: [3,80,0,4,1,41]
s: 80
i: [1,40,0,4,1,1,0,0] t33_projpl_1
a: [3,85,0,4,1,40]
s: 85
i: [1,45,0,7,1,1,0,0] t33_projpl_1
a: [3,60,0,7,1,45]
s: 60
i: [1,43,0,7,1,1,0,0] t33_projpl_1
a: [3,70,0,7,1,43]
s: 70
i: [1,54,0,2,1,0,0,2] rc2_xboole_0
a: [3,120,0,2,1,54]
s: 120
i: [1,55,0,2,1,0,0,2] rc1_xboole_0
a: [3,125,0,2,1,55]
s: 125
i: [1,56,0,2,1,0,0,1] existence_l1_incsp_1
a: [3,130,0,2,1,56]
s: 130
i: [1,52,0,4,1,0,0,1] existence_m1_subset_1
s: 52
i: [1,46,0,5,2,0,0,2] dt_u2_incsp_1
s: 46
i: [1,47,0,5,2,0,0,1] dt_u1_incsp_1
s: 47
i: [1,53,0,27,9,0,0,1] dt_k1_projpl_1
s: 53
i: [1,48,0,34,11,0,0,1] t16_projpl_1
s: 48
i: [1,49,0,34,11,0,0,1] t16_projpl_1
s: 49
i: [1,50,0,36,11,0,0,1] t16_projpl_1
s: 50
i: [1,51,0,41,13,0,0,1] t16_projpl_1
s: 51
d: [2,147,1,39,13,1,44,60,51]
d: [2,162,1,36,12,2,21,147,65]
d: [2,163,1,32,11,3,21,162,75]
d: [2,164,1,28,10,4,21,163,85]
d: [2,165,1,26,9,5,21,164,90]
d: [2,166,1,24,8,6,21,165,95]
d: [2,167,1,22,7,7,21,166,100]
d: [2,168,1,20,6,8,21,167,105]
d: [2,169,1,18,5,9,21,168,110]
d: [2,170,1,16,4,10,21,169,115]
s: 170
d: [2,148,1,39,13,1,44,70,51]
d: [2,171,1,35,12,2,21,148,80]
d: [2,172,1,31,11,3,21,171,85]
d: [2,173,1,29,10,4,21,172,90]
d: [2,174,1,27,9,5,21,173,95]
d: [2,175,1,25,8,6,21,174,100]
d: [2,176,1,23,7,7,21,175,105]
d: [2,177,1,21,6,8,21,176,110]
d: [2,178,1,19,5,9,21,177,115]
a: [3,185,1,16,4,178]
s: 185
d: [2,228,1,12,3,11,21,185,170]
s: 228
r: 170
d: [2,188,2,48,13,10,40,170,48]
d: [2,203,2,44,12,10,21,188,75]
d: [2,204,2,42,11,10,21,203,90]
d: [2,205,2,40,10,10,21,204,95]
d: [2,206,2,38,9,10,21,205,100]
d: [2,207,2,36,8,10,21,206,105]
d: [2,208,2,34,7,10,21,207,110]
d: [2,209,2,32,6,10,21,208,115]
a: [3,212,2,4,1,209]
s: 212
r: 228
d: [2,259,3,0,0,1,40,212,52]
e: 259
d: [2,253,3,0,0,1,52,52,212]
e: 253
d: [2,254,3,0,0,8,52,115,110,105,100,95,90,75,75,212,53]
e: 254
d: [2,255,3,0,0,9,52,115,110,105,100,95,90,75,80,212,53]
e: 255
d: [2,256,3,0,0,9,52,115,110,105,100,95,90,75,85,212,53]
e: 256
d: [2,257,3,0,0,8,52,115,110,105,100,95,90,75,52,212,53]
e: 257
r: 212
s: 228
d: [2,187,2,21,3,11,40,170,70]
a: [3,197,2,7,1,187]
s: 197
a: [3,182,1,3,1,178]
s: 182
r: 197
r: 228
r: 70
r: 60
r: 80
r: 182
s: 80
a: [3,181,1,3,1,178]
s: 181
s: 60
s: 70
s: 228
s: 185
a: [3,193,2,7,1,187]
s: 193
d: [2,362,3,0,0,9,52,115,110,105,100,95,90,85,80,193,53]
e: 362
d: [2,363,3,20,8,1,40,193,53]
d: [2,367,3,16,7,2,21,363,80]
d: [2,368,3,12,6,3,21,367,85]
d: [2,369,3,10,5,4,21,368,90]
d: [2,370,3,8,4,5,21,369,95]
d: [2,371,3,6,3,6,21,370,100]
d: [2,372,3,4,2,7,21,371,105]
d: [2,373,3,2,1,8,21,372,110]
d: [2,374,3,0,0,9,21,373,115]
e: 374
r: 193
d: [2,323,1,39,13,1,44,60,51]
d: [2,324,1,36,12,2,21,323,65]
d: [2,325,1,32,11,3,21,324,75]
d: [2,326,1,28,10,4,21,325,85]
d: [2,327,1,26,9,5,21,326,90]
d: [2,328,1,24,8,6,21,327,95]
d: [2,329,1,22,7,7,21,328,100]
d: [2,330,1,20,6,8,21,329,105]
d: [2,331,1,18,5,9,21,330,110]
d: [2,332,1,16,4,10,21,331,115]
s: 332
a: [3,192,2,7,1,187]
s: 192
d: [2,144,1,73,25,0,44,51,51]
d: [2,161,1,50,16,0,18,144]
s: 161
a: [3,196,2,7,1,187]
s: 196
a: [3,201,2,7,1,187]
s: 201
d: [2,402,3,0,0,11,52,115,110,105,100,95,90,192,181,85,80,201,49]
e: 402
r: 181
r: 196
r: 201
d: [2,304,0,7,1,2,23,70,182]
a: [3,316,0,7,1,304]
s: 316
s: 182
r: 60
r: 70
r: 80
r: 192
r: 228
d: [2,305,0,7,1,2,23,60,182]
a: [3,321,0,7,1,305]
s: 321
d: [2,442,1,0,0,10,52,115,110,105,100,95,90,52,65,85,75,321,48]
e: 442
d: [2,290,3,27,10,1,40,197,48]
d: [2,293,3,23,9,2,21,290,80]
d: [2,294,3,19,8,3,21,293,85]
d: [2,295,3,17,7,4,21,294,90]
d: [2,296,3,15,6,5,21,295,95]
d: [2,297,3,13,5,6,21,296,100]
d: [2,298,3,11,4,7,21,297,105]
d: [2,299,3,9,3,8,21,298,110]
d: [2,300,3,7,2,9,21,299,115]
d: [2,434,0,7,1,2,23,60,182]
f: 442,434,402,374,300,253,187,34,35,36,37,38,39,40,41,42,43,44,45
% Refutation found. Thanks to Tanya!
% SZS status Theorem for projpl_1__t33_projpl_1
% SZS output start Proof for projpl_1__t33_projpl_1
1. ! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (r1_incsp_1(X0,X3,k1_projpl_1(X0,X1,X2)) => (r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X3)) | X1 = X3)))))) [input t33_projpl_1]
2. ~! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (r1_incsp_1(X0,X3,k1_projpl_1(X0,X1,X2)) => (r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X3)) | X1 = X3)))))) [negated conjecture 1]
3. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_incsp_1(X0)) & m1_subset_1(X1,u1_incsp_1(X0)) & l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => m1_subset_1(k1_projpl_1(X0,X1,X2),u2_incsp_1(X0))) [input dt_k1_projpl_1]
9. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
12. ! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => (X1 != X2 => (((r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,X1,X3)) => k1_projpl_1(X0,X1,X2) = X3) & k1_projpl_1(X0,X1,X2) = k1_projpl_1(X0,X2,X1) & r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X2)) & r1_incsp_1(X0,X1,k1_projpl_1(X0,X1,X2)))))))) [input t16_projpl_1]
13. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X3)) & X1 != X3) & r1_incsp_1(X0,X3,k1_projpl_1(X0,X1,X2))) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & (l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0))) [ennf transformation 2]
14. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X3)) & X1 != X3 & r1_incsp_1(X0,X3,k1_projpl_1(X0,X1,X2)) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) [flattening 13]
17. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((((k1_projpl_1(X0,X1,X2) = X3 | (~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3))) & k1_projpl_1(X0,X1,X2) = k1_projpl_1(X0,X2,X1) & r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X2)) & r1_incsp_1(X0,X1,k1_projpl_1(X0,X1,X2))) | X1 = X2) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | (~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0))) [ennf transformation 12]
18. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((k1_projpl_1(X0,X1,X2) = X3 | ~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3)) & k1_projpl_1(X0,X1,X2) = k1_projpl_1(X0,X2,X1) & r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X2)) & r1_incsp_1(X0,X1,k1_projpl_1(X0,X1,X2))) | X1 = X2 | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0)) [flattening 17]
19. ! [X0,X1,X2] : (m1_subset_1(k1_projpl_1(X0,X1,X2),u2_incsp_1(X0)) | (~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0))) [ennf transformation 3]
20. ! [X0,X1,X2] : (m1_subset_1(k1_projpl_1(X0,X1,X2),u2_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0)) [flattening 19]
21. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X3)) & X1 != X3 & r1_incsp_1(X0,X3,k1_projpl_1(X0,X1,X2)) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r1_incsp_1(sK0,X2,k1_projpl_1(sK0,X1,X3)) & X1 != X3 & r1_incsp_1(sK0,X3,k1_projpl_1(sK0,X1,X2)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v4_incproj(sK0) & v3_incproj(sK0) & v2_incproj(sK0) & v1_incproj(sK0) & v6_incsp_1(sK0)) [choice axiom]
22. ? [X1] : (? [X2] : (? [X3] : (~r1_incsp_1(sK0,X2,k1_projpl_1(sK0,X1,X3)) & X1 != X3 & r1_incsp_1(sK0,X3,k1_projpl_1(sK0,X1,X2)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) => (? [X2] : (? [X3] : (~r1_incsp_1(sK0,X2,k1_projpl_1(sK0,sK1,X3)) & sK1 != X3 & r1_incsp_1(sK0,X3,k1_projpl_1(sK0,sK1,X2)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) [choice axiom]
23. ? [X2] : (? [X3] : (~r1_incsp_1(sK0,X2,k1_projpl_1(sK0,sK1,X3)) & sK1 != X3 & r1_incsp_1(sK0,X3,k1_projpl_1(sK0,sK1,X2)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) => (? [X3] : (~r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,X3)) & sK1 != X3 & r1_incsp_1(sK0,X3,k1_projpl_1(sK0,sK1,sK2)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) [choice axiom]
24. ? [X3] : (~r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,X3)) & sK1 != X3 & r1_incsp_1(sK0,X3,k1_projpl_1(sK0,sK1,sK2)) & m1_subset_1(X3,u1_incsp_1(sK0))) => (~r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,sK3)) & sK1 != sK3 & r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) & m1_subset_1(sK3,u1_incsp_1(sK0))) [choice axiom]
25. (((~r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,sK3)) & sK1 != sK3 & r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) & m1_subset_1(sK3,u1_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v4_incproj(sK0) & v3_incproj(sK0) & v2_incproj(sK0) & v1_incproj(sK0) & v6_incsp_1(sK0) [skolemisation 14,24,23,22,21]
26. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK4(X0),X0)) [choice axiom]
27. ! [X0] : m1_subset_1(sK4(X0),X0) [skolemisation 9,26]
34. v6_incsp_1(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
35. v1_incproj(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
36. v2_incproj(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
37. v3_incproj(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
38. v4_incproj(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
39. l1_incsp_1(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
40. m1_subset_1(sK1,u1_incsp_1(sK0)) [cnf transformation 25] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
41. m1_subset_1(sK2,u1_incsp_1(sK0)) [cnf transformation 25] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
42. m1_subset_1(sK3,u1_incsp_1(sK0)) [cnf transformation 25] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
43. r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
44. sK1 != sK3 [cnf transformation 25] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
45. ~r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,sK3)) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
48. r1_incsp_1(X0,X1,k1_projpl_1(X0,X1,X2)) | X1 = X2 | ~m1_subset_1(X3,u2_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0) [cnf transformation 18] {a:0,w:34,wCS:13600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X2)) | X1 = X2 | ~m1_subset_1(X3,u2_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0) [cnf transformation 18] {a:0,w:34,wCS:13600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k1_projpl_1(X0,X1,X2) = X3 | ~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3) | X1 = X2 | ~m1_subset_1(X3,u2_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0) [cnf transformation 18] {a:0,w:41,wCS:16400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. m1_subset_1(sK4(X0),X0) [cnf transformation 27] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
53. m1_subset_1(k1_projpl_1(X0,X1,X2),u2_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0) [cnf transformation 20] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. 1 <=> r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,sK3)) [avatar definition]
60. ~r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,sK3)) <- (~1) [avatar component clause 58] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
61. ~1 [avatar split clause 45,58]
63. 2 <=> sK1 = sK3 [avatar definition]
65. sK1 != sK3 <- (~2) [avatar component clause 63] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
66. ~2 [avatar split clause 44,63]
68. 3 <=> r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) [avatar definition]
70. r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) <- (3) [avatar component clause 68] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
71. 3 [avatar split clause 43,68]
73. 4 <=> m1_subset_1(sK3,u1_incsp_1(sK0)) [avatar definition]
75. m1_subset_1(sK3,u1_incsp_1(sK0)) <- (4) [avatar component clause 73] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
76. 4 [avatar split clause 42,73]
78. 5 <=> m1_subset_1(sK2,u1_incsp_1(sK0)) [avatar definition]
80. m1_subset_1(sK2,u1_incsp_1(sK0)) <- (5) [avatar component clause 78] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
81. 5 [avatar split clause 41,78]
83. 6 <=> m1_subset_1(sK1,u1_incsp_1(sK0)) [avatar definition]
85. m1_subset_1(sK1,u1_incsp_1(sK0)) <- (6) [avatar component clause 83] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
86. 6 [avatar split clause 40,83]
88. 7 <=> l1_incsp_1(sK0) [avatar definition]
90. l1_incsp_1(sK0) <- (7) [avatar component clause 88] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
91. 7 [avatar split clause 39,88]
93. 8 <=> v4_incproj(sK0) [avatar definition]
95. v4_incproj(sK0) <- (8) [avatar component clause 93] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
96. 8 [avatar split clause 38,93]
98. 9 <=> v3_incproj(sK0) [avatar definition]
100. v3_incproj(sK0) <- (9) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
101. 9 [avatar split clause 37,98]
103. 10 <=> v2_incproj(sK0) [avatar definition]
105. v2_incproj(sK0) <- (10) [avatar component clause 103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
106. 10 [avatar split clause 36,103]
108. 11 <=> v1_incproj(sK0) [avatar definition]
110. v1_incproj(sK0) <- (11) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
111. 11 [avatar split clause 35,108]
113. 12 <=> v6_incsp_1(sK0) [avatar definition]
115. v6_incsp_1(sK0) <- (12) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
116. 12 [avatar split clause 34,113]
147. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | sK1 = sK3 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1) [superposition 60,51] {a:1,w:39,wCS:3900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
162. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~2) [subsumption resolution 147,65] {a:1,w:36,wCS:3600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
163. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~2, 4) [subsumption resolution 162,75] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
164. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~2, 4, 6) [subsumption resolution 163,85] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
165. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~2, 4, 6, 7) [subsumption resolution 164,90] {a:1,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
166. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~2, 4, 6, 7, 8) [subsumption resolution 165,95] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
167. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~2, 4, 6, 7, 8, 9) [subsumption resolution 166,100] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
168. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~2, 4, 6, 7, 8, 9, 10) [subsumption resolution 167,105] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
169. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v6_incsp_1(sK0) <- (~1, ~2, 4, 6, 7, 8, 9, 10, 11) [subsumption resolution 168,110] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
170. ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) <- (~1, ~2, 4, 6, 7, 8, 9, 10, 11, 12) [subsumption resolution 169,115] {a:1,w:16,wCS:1600,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
180. 16 <=> sK1 = sK2 [avatar definition]
181. sK1 != sK2 <- (~16) [avatar component clause 180] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
182. sK1 = sK2 <- (16) [avatar component clause 180] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
187. ~r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,sK2)) | ~r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(k1_projpl_1(sK0,sK1,sK2),u2_incsp_1(sK0)) <- (~1, ~2, 3, 4, 6, 7, 8, 9, 10, 11, 12) [resolution 170,70] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
191. 18 <=> m1_subset_1(k1_projpl_1(sK0,sK1,sK2),u2_incsp_1(sK0)) [avatar definition]
192. m1_subset_1(k1_projpl_1(sK0,sK1,sK2),u2_incsp_1(sK0)) <- (18) [avatar component clause 191] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
193. ~m1_subset_1(k1_projpl_1(sK0,sK1,sK2),u2_incsp_1(sK0)) <- (~18) [avatar component clause 191] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
195. 19 <=> r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK1,sK2)) [avatar definition]
197. ~r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK1,sK2)) <- (~19) [avatar component clause 195] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
199. 20 <=> r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,sK2)) [avatar definition]
201. ~r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK1,sK2)) <- (~20) [avatar component clause 199] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
202. ~18 | ~19 | ~20 | 1 | 2 | ~3 | ~4 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 [avatar split clause 187,113,108,103,98,93,88,83,73,68,63,58,199,195,191]
211. 21 <=> ! [X1] : ~m1_subset_1(X1,u2_incsp_1(sK0)) [avatar definition]
212. ~m1_subset_1(X1,u2_incsp_1(sK0)) <- (21) [avatar component clause 211] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
253. $false <- (21) [unit resulting resolution 52,212] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
261. ~21 [avatar contradiction clause 253]
290. sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~19) [resolution 197,48] {a:3,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
293. sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, ~19) [subsumption resolution 290,80] {a:3,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
294. sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, ~19) [subsumption resolution 293,85] {a:3,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
295. sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, 7, ~19) [subsumption resolution 294,90] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
296. sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, 7, 8, ~19) [subsumption resolution 295,95] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
297. sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, 7, 8, 9, ~19) [subsumption resolution 296,100] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
298. sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, 7, 8, 9, 10, ~19) [subsumption resolution 297,105] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
299. sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v6_incsp_1(sK0) <- (5, 6, 7, 8, 9, 10, 11, ~19) [subsumption resolution 298,110] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
300. sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) <- (5, 6, 7, 8, 9, 10, 11, 12, ~19) [subsumption resolution 299,115] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
301. 21 | 16 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 19 [avatar split clause 300,195,113,108,103,98,93,88,83,78,180,211]
319. 28 <=> r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK1,sK3)) [avatar definition]
321. ~r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK1,sK3)) <- (~28) [avatar component clause 319] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
363. ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~18) [resolution 193,53] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
367. ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, ~18) [subsumption resolution 363,80] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
368. ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, ~18) [subsumption resolution 367,85] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
369. ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, 7, ~18) [subsumption resolution 368,90] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
370. ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, 7, 8, ~18) [subsumption resolution 369,95] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
371. ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, 7, 8, 9, ~18) [subsumption resolution 370,100] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
372. ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (5, 6, 7, 8, 9, 10, ~18) [subsumption resolution 371,105] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
373. ~v6_incsp_1(sK0) <- (5, 6, 7, 8, 9, 10, 11, ~18) [subsumption resolution 372,110] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
374. $false <- (5, 6, 7, 8, 9, 10, 11, 12, ~18) [subsumption resolution 373,115] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
375. ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 18 [avatar contradiction clause 374]
402. $false <- (5, 6, 7, 8, 9, 10, 11, 12, ~16, 18, ~20) [unit resulting resolution 115,110,105,100,95,90,192,181,85,80,201,49] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:69,thDist:-69}
406. ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 16 | ~18 | 20 [avatar contradiction clause 402]
434. ~r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK1,sK3)) <- (~1, 16) [backward demodulation 60,182] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
440. ~28 | 1 | ~16 [avatar split clause 434,180,58,319]
442. $false <- (~2, 4, 6, 7, 8, 9, 10, 11, 12, ~28) [unit resulting resolution 115,110,105,100,95,90,52,65,85,75,321,48] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:53,thDist:-53}
446. 2 | ~4 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 28 [avatar contradiction clause 442]
465. $false [avatar sat refutation 61,66,71,76,81,86,91,96,101,106,111,116,202,261,301,375,406,440,446] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:790,thDist:-790}
% SZS output end Proof for projpl_1__t33_projpl_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 315
% Active clauses: 42
% Passive clauses: 59
% Final active clauses: 24
% Final passive clauses: 12
% Main loop iterations started: 45
% 
% >>> Simplifying Inferences
% Duplicate literals: 143
% Fw subsumption resolutions: 166
% Bw demodulations: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 19
% Backward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 8
% Forward superposition: 10
% Backward superposition: 13
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 25
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 46
% SAT solver unit clauses: 17
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

