i: [1,290,0,7,2,1,0,0] t5_extreal2
a: [3,446,0,3,1,290]
s: 446
i: [1,288,0,3,1,1,0,0] t5_extreal2
a: [3,456,0,3,1,288]
s: 456
a: [3,449,0,4,1,290]
s: 449
i: [1,291,0,2,1,0,0,2] fc1_numbers
a: [3,461,0,2,1,291]
s: 461
i: [1,292,0,2,1,0,0,1] fc5_numbers
a: [3,466,0,2,1,292]
s: 466
i: [1,293,0,2,1,0,0,255] fc2_membered
a: [3,471,0,2,1,293]
s: 471
i: [1,294,0,2,1,0,0,255] fc3_membered
a: [3,476,0,2,1,294]
s: 476
i: [1,295,0,2,1,0,0,255] fc6_membered
a: [3,481,0,2,1,295]
s: 481
i: [1,296,0,2,1,0,0,255] fc1_xboole_0
a: [3,486,0,2,1,296]
s: 486
i: [1,297,0,2,1,0,0,255] spc0_boole
a: [3,491,0,2,1,297]
s: 491
i: [1,406,0,7,3,0,0,2] reflexivity_r1_xxreal_0
a: [3,545,0,2,1,406]
s: 545
i: [1,408,0,2,1,0,0,2] rc2_xboole_0
a: [3,553,0,2,1,408]
s: 553
i: [1,409,0,2,1,0,0,3] rc1_xreal_0
a: [3,558,0,2,1,409]
s: 558
i: [1,410,0,2,1,0,0,2] rc1_xboole_0
a: [3,563,0,2,1,410]
s: 563
i: [1,411,0,2,1,0,0,2] rc1_xxreal_0
a: [3,568,0,2,1,411]
s: 568
d: [2,720,0,0,0,2,21,568,545]
e: 720
r: 545
s: 568
i: [1,413,0,2,1,0,0,3] rc1_membered
a: [3,573,0,2,1,413]
s: 573
i: [1,412,0,2,1,0,0,3] rc1_membered
a: [3,578,0,2,1,412]
s: 578
i: [1,415,0,2,1,0,0,3] rc2_membered
a: [3,583,0,2,1,415]
s: 583
i: [1,414,0,2,1,0,0,3] rc2_membered
a: [3,588,0,2,1,414]
s: 588
i: [1,417,0,2,1,0,0,255] rc3_membered
a: [3,593,0,2,1,417]
s: 593
i: [1,416,0,2,1,0,0,255] rc3_membered
a: [3,598,0,2,1,416]
s: 598
i: [1,421,0,2,1,0,0,255] rc2_xreal_0
a: [3,603,0,2,1,421]
s: 603
i: [1,420,0,2,1,0,0,255] rc2_xreal_0
a: [3,608,0,2,1,420]
s: 608
i: [1,419,0,2,1,0,0,255] rc2_xreal_0
a: [3,613,0,2,1,419]
s: 613
i: [1,418,0,2,1,0,0,255] rc2_xreal_0
a: [3,618,0,2,1,418]
s: 618
i: [1,425,0,2,1,0,0,255] rc3_xreal_0
a: [3,623,0,2,1,425]
s: 623
i: [1,424,0,2,1,0,0,255] rc3_xreal_0
a: [3,628,0,2,1,424]
s: 628
i: [1,423,0,2,1,0,0,255] rc3_xreal_0
a: [3,633,0,2,1,423]
s: 633
i: [1,422,0,2,1,0,0,255] rc3_xreal_0
a: [3,638,0,2,1,422]
s: 638
i: [1,429,0,2,1,0,0,255] rc4_xreal_0
a: [3,643,0,2,1,429]
s: 643
i: [1,428,0,2,1,0,0,255] rc4_xreal_0
a: [3,648,0,2,1,428]
s: 648
i: [1,427,0,2,1,0,0,255] rc4_xreal_0
a: [3,653,0,2,1,427]
s: 653
i: [1,426,0,2,1,0,0,255] rc4_xreal_0
a: [3,658,0,2,1,426]
s: 658
i: [1,431,0,2,1,0,0,2] rc4_xxreal_0
a: [3,663,0,2,1,431]
s: 663
i: [1,430,0,2,1,0,0,2] rc4_xxreal_0
a: [3,668,0,2,1,430]
s: 668
i: [1,433,0,2,1,0,0,3] rc2_xxreal_0
a: [3,673,0,2,1,433]
s: 673
i: [1,432,0,2,1,0,0,3] rc2_xxreal_0
a: [3,678,0,2,1,432]
s: 678
i: [1,436,0,2,1,0,0,3] rc1_xxreal_3
a: [3,683,0,2,1,436]
s: 683
i: [1,435,0,2,1,0,0,3] rc1_xxreal_3
a: [3,688,0,2,1,435]
s: 688
i: [1,434,0,2,1,0,0,3] rc1_xxreal_3
a: [3,693,0,2,1,434]
s: 693
i: [1,438,0,2,1,0,0,2] rc3_xxreal_0
a: [3,698,0,2,1,438]
s: 698
i: [1,437,0,2,1,0,0,2] rc3_xxreal_0
a: [3,703,0,2,1,437]
s: 703
i: [1,441,0,2,1,0,0,2] rc2_xxreal_3
a: [3,708,0,2,1,441]
s: 708
i: [1,440,0,2,1,0,0,2] rc2_xxreal_3
a: [3,713,0,2,1,440]
s: 713
i: [1,439,0,2,1,0,0,2] rc2_xxreal_3
a: [3,718,0,2,1,439]
s: 718
i: [1,399,0,3,1,0,0,255] reflexivity_r1_tarski
s: 399
i: [1,298,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,496,0,3,1,298]
s: 496
r: 481
d: [2,722,0,2,1,2,23,481,496]
a: [3,726,0,2,1,722]
s: 726
i: [1,299,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,501,0,3,1,299]
s: 501
r: 486
d: [2,728,0,2,1,2,23,486,501]
a: [3,732,0,2,1,728]
s: 732
i: [1,300,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r0_r0
a: [3,506,0,3,1,300]
s: 506
i: [1,301,0,3,1,0,0,255] t1_numerals
a: [3,511,0,3,1,301]
s: 511
i: [1,307,0,3,1,0,0,255] spc0_numerals
a: [3,531,0,3,1,307]
s: 531
i: [1,306,0,3,1,0,0,255] spc0_numerals
a: [3,536,0,3,1,306]
s: 536
d: [2,734,0,3,1,2,22,511,496]
d: [2,735,0,3,1,3,22,734,501]
a: [3,739,0,3,1,735]
s: 739
i: [1,308,0,4,2,0,0,255] cc2_xreal_0
s: 308
i: [1,309,0,4,2,0,0,255] cc2_xxreal_0
s: 309
i: [1,312,0,4,2,0,0,255] cc2_membered
s: 312
i: [1,313,0,4,2,0,0,255] cc3_membered
s: 313
i: [1,318,0,4,2,0,0,255] cc5_membered
s: 318
d: [2,741,1,2,1,1,40,318,476]
a: [3,745,1,2,1,741]
s: 745
i: [1,319,0,4,2,0,0,255] cc4_membered
s: 319
d: [2,747,1,2,1,1,40,319,476]
a: [3,751,1,2,1,747]
s: 751
i: [1,322,0,4,2,0,0,255] cc1_membered
s: 322
i: [1,327,0,4,2,0,0,255] cc3_xreal_0
s: 327
d: [2,753,1,2,1,1,40,322,726]
a: [3,760,1,2,1,753]
s: 760
i: [1,328,0,4,2,0,0,3] cc4_xreal_0
s: 328
d: [2,754,1,2,1,1,40,322,573]
a: [3,765,1,2,1,754]
s: 765
i: [1,341,0,4,2,0,0,3] cc18_membered
s: 341
i: [1,398,0,4,1,0,0,1] existence_m1_subset_1
s: 398
d: [2,755,1,2,1,1,40,322,583]
a: [3,770,1,2,1,755]
s: 770
i: [1,302,0,4,1,0,0,255] rqRealNeg__k4_xcmplx_0__r0_r0
a: [3,516,0,4,1,302]
s: 516
d: [2,756,1,2,1,1,40,322,593]
a: [3,775,1,2,1,756]
s: 775
i: [1,303,0,4,1,0,0,2] dt_k5_numbers
a: [3,521,0,4,1,303]
s: 521
i: [1,304,0,4,1,0,0,1] dt_k6_numbers
a: [3,526,0,4,1,304]
s: 526
d: [2,777,1,2,1,1,40,327,558]
a: [3,784,1,2,1,777]
s: 784
i: [1,305,0,4,1,0,0,255] spc0_numerals
a: [3,541,0,4,1,305]
s: 541
d: [2,792,1,2,1,1,40,328,558]
a: [3,799,1,2,1,792]
s: 799
i: [1,325,0,5,2,0,0,255] dt_k4_xcmplx_0
s: 325
i: [1,329,0,5,2,0,0,255] fc3_xreal_0
s: 329
d: [2,807,1,2,1,1,40,341,732]
a: [3,815,1,2,1,807]
s: 815
i: [1,330,0,5,2,0,0,255] fc3_xreal_0
s: 330
d: [2,808,1,2,1,1,40,341,491]
a: [3,820,1,2,1,808]
s: 820
i: [1,331,0,5,2,0,0,255] fc1_xxreal_3
s: 331
i: [1,332,0,5,2,0,0,255] fc1_xxreal_3
s: 332
d: [2,809,1,2,1,1,40,341,563]
a: [3,825,1,2,1,809]
s: 825
i: [1,342,0,5,2,0,0,255] t6_boole
s: 342
d: [2,810,1,2,1,1,40,341,658]
a: [3,830,1,2,1,810]
s: 830
i: [1,343,0,5,2,0,0,255] dt_k2_xxreal_3
s: 343
i: [1,345,0,5,2,0,0,2] cc25_membered
s: 345
d: [2,811,1,2,1,1,40,341,668]
a: [3,835,1,2,1,811]
s: 835
i: [1,346,0,5,2,0,0,2] cc1_xreal_0
s: 346
d: [2,895,1,2,1,1,40,346,531]
a: [3,900,1,2,1,895]
s: 900
i: [1,347,0,5,2,0,0,1] cc1_xxreal_0
s: 347
i: [1,395,0,5,2,0,0,255] t7_boole
s: 395
d: [2,919,1,2,1,1,40,347,456]
a: [3,924,1,2,1,919]
s: 924
a: [3,548,0,5,2,406]
s: 548
d: [2,786,2,2,1,1,40,760,312]
a: [3,790,2,2,1,786]
s: 790
d: [2,874,0,5,2,1,22,342,501]
s: 874
i: [1,348,0,6,2,0,0,1] fc1_extreal1
s: 348
d: [2,942,1,5,2,1,44,348,449]
d: [2,943,1,2,1,2,21,942,456]
a: [3,947,1,2,1,943]
s: 947
i: [1,357,0,6,2,0,0,2] cc11_membered
s: 357
d: [2,949,2,4,2,2,44,947,874]
s: 949
i: [1,358,0,6,2,0,0,255] cc8_membered
s: 358
i: [1,359,0,6,2,0,0,255] cc7_membered
s: 359
d: [2,957,3,2,1,3,40,949,491]
a: [3,964,3,2,1,957]
s: 964
i: [1,369,0,6,3,0,0,3] cc7_xxreal_0
s: 369
d: [2,958,3,2,1,3,40,949,563]
a: [3,969,3,2,1,958]
s: 969
i: [1,370,0,6,3,0,0,3] cc7_xxreal_0
s: 370
i: [1,381,0,6,3,0,0,3] cc3_xxreal_0
s: 381
d: [2,959,3,2,1,3,40,949,658]
a: [3,974,3,2,1,959]
s: 974
i: [1,383,0,6,3,0,0,3] cc3_xxreal_0
s: 383
d: [2,1037,2,4,2,1,40,383,924]
a: [3,1051,2,2,1,1037]
s: 1051
i: [1,386,0,6,3,0,0,3] cc5_xxreal_0
s: 386
i: [1,388,0,6,3,0,0,3] cc5_xxreal_0
s: 388
d: [2,960,3,2,1,3,40,949,668]
a: [3,979,3,2,1,960]
s: 979
i: [1,391,0,6,2,0,0,255] t1_subset
s: 391
d: [2,995,1,4,2,1,40,369,732]
a: [3,1003,1,2,1,995]
s: 1003
i: [1,400,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 400
i: [1,310,0,7,3,0,0,255] cc12_membered
s: 310
d: [2,1117,1,4,2,1,40,310,456]
a: [3,1143,1,2,1,1117]
s: 1143
i: [1,316,0,7,3,0,0,255] cc13_membered
s: 316
d: [2,996,1,4,2,1,40,369,491]
a: [3,1012,1,2,1,996]
s: 1012
i: [1,320,0,7,3,0,0,255] cc14_membered
s: 320
i: [1,323,0,7,3,0,0,255] cc17_membered
s: 323
d: [2,1176,1,4,2,1,40,320,456]
a: [3,1195,1,2,1,1176]
s: 1195
i: [1,326,0,7,2,0,0,255] involutiveness_k4_xcmplx_0
s: 326
d: [2,1205,1,4,2,1,40,323,456]
a: [3,1236,1,2,1,1205]
s: 1236
i: [1,344,0,7,2,0,0,255] involutiveness_k2_xxreal_3
s: 344
i: [1,349,0,7,2,0,0,255] dt_k2_supinf_2
s: 349
d: [2,997,1,4,2,1,40,369,563]
a: [3,1021,1,2,1,997]
s: 1021
i: [1,350,0,7,2,0,0,1] dt_k3_extreal1
s: 350
d: [2,1260,1,6,2,1,44,350,449]
d: [2,1262,1,3,1,2,21,1260,456]
a: [3,1266,1,3,1,1262]
s: 1266
i: [1,360,0,7,3,0,0,255] fc15_xreal_0
s: 360
i: [1,361,0,7,3,0,0,255] fc15_xreal_0
s: 361
d: [2,1273,2,5,2,2,44,1266,874]
s: 1273
i: [1,362,0,7,3,0,0,255] fc16_xreal_0
s: 362
d: [2,1279,3,4,2,2,40,1273,347]
s: 1279
d: [2,1285,4,0,0,4,52,1021,563,1279]
e: 1285
d: [2,1288,4,2,1,3,40,1279,563]
d: [2,1292,4,0,0,4,21,1288,1021]
e: 1292
r: 1021
i: [1,363,0,7,3,0,0,255] fc16_xreal_0
s: 363
i: [1,375,0,7,3,0,0,255] fc14_xxreal_3
s: 375
d: [2,1282,3,6,3,2,40,1273,316]
d: [2,1284,3,4,2,3,21,1282,471]
s: 1284
i: [1,376,0,7,3,0,0,255] fc14_xxreal_3
s: 376
a: [3,1006,1,2,1,995]
s: 1006
i: [1,379,0,7,3,0,0,255] fc15_xxreal_3
s: 379
i: [1,380,0,7,3,0,0,255] fc15_xxreal_3
s: 380
a: [3,1017,1,2,1,997]
s: 1017
i: [1,384,0,7,3,0,0,255] fc16_xxreal_3
s: 384
a: [3,1020,1,2,1,997]
s: 1020
i: [1,385,0,7,3,0,0,255] fc16_xxreal_3
s: 385
i: [1,389,0,7,3,0,0,255] fc17_xxreal_3
s: 389
d: [2,998,1,4,2,1,40,369,658]
d: [2,1023,1,2,1,2,21,998,648]
a: [3,1027,1,2,1,1023]
s: 1027
i: [1,390,0,7,3,0,0,255] fc17_xxreal_3
s: 390
d: [2,999,1,4,2,1,40,369,668]
d: [2,1029,1,2,1,2,21,999,663]
a: [3,1033,1,2,1,1029]
s: 1033
i: [1,393,0,7,2,0,0,255] t3_subset
s: 393
i: [1,394,0,7,3,0,0,1] t8_boole
s: 394
d: [2,1039,1,4,2,1,40,383,568]
a: [3,1073,1,2,1,1039]
s: 1073
i: [1,311,0,8,3,0,0,255] cc19_membered
s: 311
d: [2,1040,1,4,2,1,40,383,613]
d: [2,1075,1,2,1,2,21,1040,608]
a: [3,1079,1,2,1,1075]
s: 1079
i: [1,314,0,8,3,0,0,255] cc23_membered
s: 314
i: [1,315,0,8,3,0,0,255] cc22_membered
s: 315
d: [2,1041,1,4,2,1,40,383,633]
d: [2,1081,1,2,1,2,21,1041,628]
a: [3,1085,1,2,1,1081]
s: 1085
i: [1,317,0,8,3,0,0,255] cc20_membered
s: 317
d: [2,1044,1,4,2,1,40,383,678]
d: [2,1088,1,2,1,2,21,1044,673]
a: [3,1092,1,2,1,1088]
s: 1092
i: [1,321,0,8,3,0,0,255] cc21_membered
s: 321
i: [1,324,0,8,3,0,0,3] cc24_membered
s: 324
d: [2,1045,1,4,2,1,40,383,693]
d: [2,1094,1,2,1,2,21,1045,688]
a: [3,1098,1,2,1,1094]
s: 1098
i: [1,351,0,8,2,0,0,255] redefinition_k2_supinf_2
s: 351
d: [2,1046,1,4,2,1,40,383,703]
d: [2,1100,1,2,1,2,21,1046,698]
a: [3,1104,1,2,1,1100]
s: 1104
i: [1,352,0,8,2,0,0,255] involutiveness_k2_supinf_2
s: 352
i: [1,365,0,8,4,0,0,3] cc6_xxreal_0
s: 365
a: [3,1054,2,2,1,1037]
s: 1054
i: [1,367,0,8,4,0,0,3] cc4_xxreal_0
s: 367
a: [3,1068,1,2,1,1039]
s: 1068
i: [1,372,0,8,4,0,0,3] cc2_xxreal_3
s: 372
r: 446
r: 449
r: 949
r: 947
r: 969
r: 979
r: 1051
r: 1054
a: [3,445,0,3,1,290]
s: 445
r: 924
r: 456
a: [3,946,1,2,1,943]
s: 946
a: [3,450,0,4,1,290]
s: 450
a: [3,1050,2,2,1,1037]
s: 1050
d: [2,1418,0,4,2,2,21,386,1279]
s: 1418
d: [2,1434,2,0,0,4,52,946,732,1418]
e: 1434
d: [2,1435,1,2,1,3,40,1418,732]
d: [2,1441,1,0,0,4,21,1435,946]
e: 1441
r: 1418
r: 1279
r: 1273
r: 1266
r: 946
d: [2,1427,0,4,1,2,22,450,445]
a: [3,1431,0,4,1,1427]
s: 1431
s: 947
s: 370
d: [2,1394,1,6,3,1,40,372,663]
d: [2,1409,1,4,2,2,21,1394,1033]
d: [2,1410,1,2,1,3,21,1409,979]
a: [3,1414,1,2,1,1410]
s: 1414
s: 386
i: [1,373,0,8,4,0,0,3] cc8_xxreal_0
s: 373
d: [2,1389,2,6,3,1,40,372,1020]
d: [2,1402,2,4,2,2,21,1389,1017]
d: [2,1403,2,2,1,3,21,1402,969]
a: [3,1407,2,2,1,1403]
s: 1407
i: [1,378,0,8,4,0,0,3] cc1_xxreal_3
s: 378
d: [2,1463,2,2,1,1,40,1414,327]
a: [3,1467,2,2,1,1463]
s: 1467
i: [1,392,0,8,3,0,0,255] t2_subset
s: 392
i: [1,353,0,9,2,0,0,1] projectivity_k3_extreal1
s: 353
s: 969
i: [1,397,0,9,3,0,0,255] t5_subset
s: 397
s: 979
i: [1,442,0,9,3,0,0,255] ie2_xxreal_3
s: 442
d: [2,1521,0,7,2,0,21,442,327]
s: 1521
d: [2,1470,3,2,1,1,40,1407,327]
a: [3,1474,3,2,1,1470]
s: 1474
i: [1,354,0,10,3,0,0,255] d1_extreal1
s: 354
s: 1266
i: [1,356,0,10,3,0,0,1] t4_extreal2
s: 356
i: [1,396,0,10,3,0,0,255] t4_subset
s: 396
d: [2,1541,1,9,3,1,44,1431,354]
d: [2,1542,1,6,2,1,19,1541]
a: [3,1549,1,3,1,1542]
s: 1549
d: [2,1572,2,0,0,3,52,1006,1549,548]
e: 1572
d: [2,1573,2,2,1,2,40,1549,548]
d: [2,1576,2,0,0,3,21,1573,1006]
e: 1576
r: 445
r: 1431
r: 1549
s: 446
s: 924
s: 456
s: 449
i: [1,289,0,7,2,1,0,0] t5_extreal2
d: [2,1149,1,4,2,1,40,316,739]
d: [2,1331,1,4,2,1,40,317,521]
d: [2,1333,1,2,1,2,21,1331,751]
d: [2,1425,0,3,1,2,23,456,445]
d: [2,1597,1,9,3,1,44,356,449]
d: [2,1609,1,6,2,2,21,1597,446]
d: [2,1610,1,3,1,3,21,1609,456]
f: 1610,1576,1542,1425,1427,1333,1262,1149,747,735,720,411,406,303,301,299,298,294,288,289,290
% Refutation found. Thanks to Tanya!
% SZS status Theorem for extreal2__t5_extreal2
% SZS output start Proof for extreal2__t5_extreal2
1. ! [X0] : (m1_subset_1(X0,k7_numbers) => (k6_numbers = X0 <=> k6_numbers = k3_extreal1(X0))) [input t5_extreal2]
2. ~! [X0] : (m1_subset_1(X0,k7_numbers) => (k6_numbers = X0 <=> k6_numbers = k3_extreal1(X0))) [negated conjecture 1]
32. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v2_membered(X1))) [input cc20_membered]
38. ! [X0] : (v3_membered(X0) => v2_membered(X0)) [input cc4_membered]
57. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
72. k5_numbers = k4_ordinal1 [input redefinition_k5_numbers]
77. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [input dt_k5_numbers]
79. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xxreal_0(X1))) [input cc13_membered]
91. v3_membered(k1_numbers) [input fc3_membered]
95. ? [X0] : v1_xxreal_0(X0) [input rc1_xxreal_0]
105. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => r1_xxreal_0(X0,X0)) [input reflexivity_r1_xxreal_0]
109. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
111. ! [X0] : (m1_subset_1(X0,k7_numbers) => m1_subset_1(k3_extreal1(X0),k7_numbers)) [input dt_k3_extreal1]
121. ! [X0] : (m1_subset_1(X0,k7_numbers) => ~(r1_xxreal_0(k3_extreal1(X0),k6_numbers) & k6_numbers != X0)) [input t4_extreal2]
122. ! [X0] : (m1_subset_1(X0,k7_numbers) => ((~r1_xxreal_0(k6_numbers,X0) => k3_extreal1(X0) = k2_supinf_2(X0)) & (r1_xxreal_0(k6_numbers,X0) => k3_extreal1(X0) = X0))) [input d1_extreal1]
140. ? [X0] : ((k6_numbers = X0 <~> k6_numbers = k3_extreal1(X0)) & m1_subset_1(X0,k7_numbers)) [ennf transformation 2]
149. ! [X0] : (! [X1] : (v1_xxreal_0(X1) | ~m1_subset_1(X1,X0)) | ~v2_membered(X0)) [ennf transformation 79]
150. ! [X0] : (! [X1] : (v2_membered(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v2_membered(X0)) [ennf transformation 32]
152. ! [X0] : (v2_membered(X0) | ~v3_membered(X0)) [ennf transformation 38]
189. ! [X0] : (m1_subset_1(k3_extreal1(X0),k7_numbers) | ~m1_subset_1(X0,k7_numbers)) [ennf transformation 111]
193. ! [X0] : (((k3_extreal1(X0) = k2_supinf_2(X0) | r1_xxreal_0(k6_numbers,X0)) & (k3_extreal1(X0) = X0 | ~r1_xxreal_0(k6_numbers,X0))) | ~m1_subset_1(X0,k7_numbers)) [ennf transformation 122]
194. ! [X0] : ((~r1_xxreal_0(k3_extreal1(X0),k6_numbers) | k6_numbers = X0) | ~m1_subset_1(X0,k7_numbers)) [ennf transformation 121]
195. ! [X0] : (~r1_xxreal_0(k3_extreal1(X0),k6_numbers) | k6_numbers = X0 | ~m1_subset_1(X0,k7_numbers)) [flattening 194]
245. ! [X0,X1] : (r1_xxreal_0(X0,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 105]
246. ! [X0,X1] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 245]
249. ? [X0] : (((k6_numbers != k3_extreal1(X0) | k6_numbers != X0) & (k6_numbers = k3_extreal1(X0) | k6_numbers = X0)) & m1_subset_1(X0,k7_numbers)) [nnf transformation 140]
250. ? [X0] : ((k6_numbers != k3_extreal1(X0) | k6_numbers != X0) & (k6_numbers = k3_extreal1(X0) | k6_numbers = X0) & m1_subset_1(X0,k7_numbers)) [flattening 249]
251. ? [X0] : ((k6_numbers != k3_extreal1(X0) | k6_numbers != X0) & (k6_numbers = k3_extreal1(X0) | k6_numbers = X0) & m1_subset_1(X0,k7_numbers)) => ((k6_numbers != k3_extreal1(sK0) | k6_numbers != sK0) & (k6_numbers = k3_extreal1(sK0) | k6_numbers = sK0) & m1_subset_1(sK0,k7_numbers)) [choice axiom]
252. (k6_numbers != k3_extreal1(sK0) | k6_numbers != sK0) & (k6_numbers = k3_extreal1(sK0) | k6_numbers = sK0) & m1_subset_1(sK0,k7_numbers) [skolemisation 250,251]
264. ? [X0] : v1_xxreal_0(X0) => v1_xxreal_0(sK6) [choice axiom]
265. v1_xxreal_0(sK6) [skolemisation 95,264]
288. m1_subset_1(sK0,k7_numbers) [cnf transformation 252] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
289. k6_numbers = k3_extreal1(sK0) | k6_numbers = sK0 [cnf transformation 252] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
290. k6_numbers != k3_extreal1(sK0) | k6_numbers != sK0 [cnf transformation 252] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
294. v3_membered(k1_numbers) [cnf transformation 91] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
298. k5_numbers = k4_ordinal1 [cnf transformation 72] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
299. k6_numbers = k1_xboole_0 [cnf transformation 109] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
301. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 57] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
303. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [cnf transformation 77] {a:0,w:4,wCS:1600,sine:2,thAx:0,allAx:1,thDist:-1}
316. ~m1_subset_1(X1,X0) | v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 149] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
317. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v2_membered(X1) | ~v2_membered(X0) [cnf transformation 150] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
319. ~v3_membered(X0) | v2_membered(X0) [cnf transformation 152] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
350. m1_subset_1(k3_extreal1(X0),k7_numbers) | ~m1_subset_1(X0,k7_numbers) [cnf transformation 189] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
354. k3_extreal1(X0) = X0 | ~r1_xxreal_0(k6_numbers,X0) | ~m1_subset_1(X0,k7_numbers) [cnf transformation 193] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
356. ~r1_xxreal_0(k3_extreal1(X0),k6_numbers) | k6_numbers = X0 | ~m1_subset_1(X0,k7_numbers) [cnf transformation 195] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
406. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 246] {a:0,w:7,wCS:2800,sine:2,thAx:0,allAx:1,thDist:-1}
411. v1_xxreal_0(sK6) [cnf transformation 265] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
444. 1 <=> k6_numbers = sK0 [avatar definition]
445. k6_numbers = sK0 <- (1) [avatar component clause 444] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
446. k6_numbers != sK0 <- (~1) [avatar component clause 444] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
448. 2 <=> k6_numbers = k3_extreal1(sK0) [avatar definition]
449. k6_numbers = k3_extreal1(sK0) <- (2) [avatar component clause 448] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
450. k6_numbers != k3_extreal1(sK0) <- (~2) [avatar component clause 448] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
451. ~1 | ~2 [avatar split clause 290,448,444]
452. 1 | 2 [avatar split clause 289,448,444]
454. 3 <=> m1_subset_1(sK0,k7_numbers) [avatar definition]
456. m1_subset_1(sK0,k7_numbers) <- (3) [avatar component clause 454] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
457. 3 [avatar split clause 288,454]
474. 7 <=> v3_membered(k1_numbers) [avatar definition]
476. v3_membered(k1_numbers) <- (7) [avatar component clause 474] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
477. 7 [avatar split clause 294,474]
494. 11 <=> k5_numbers = k4_ordinal1 [avatar definition]
496. k5_numbers = k4_ordinal1 <- (11) [avatar component clause 494] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
497. 11 [avatar split clause 298,494]
499. 12 <=> k6_numbers = k1_xboole_0 [avatar definition]
501. k6_numbers = k1_xboole_0 <- (12) [avatar component clause 499] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
502. 12 [avatar split clause 299,499]
509. 14 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
511. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (14) [avatar component clause 509] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
512. 14 [avatar split clause 301,509]
519. 16 <=> m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [avatar definition]
521. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) <- (16) [avatar component clause 519] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
522. 16 [avatar split clause 303,519]
544. 21 <=> ! [X1] : ~v1_xxreal_0(X1) [avatar definition]
545. ~v1_xxreal_0(X1) <- (21) [avatar component clause 544] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
547. 22 <=> ! [X0] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0)) [avatar definition]
548. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) <- (22) [avatar component clause 547] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
549. 21 | 22 [avatar split clause 406,547,544]
566. 26 <=> v1_xxreal_0(sK6) [avatar definition]
568. v1_xxreal_0(sK6) <- (26) [avatar component clause 566] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
569. 26 [avatar split clause 411,566]
720. $false <- (21, 26) [subsumption resolution 568,545] {a:0,w:0,sine:255,thAx:0,allAx:2,thDist:-2}
721. ~21 | ~26 [avatar contradiction clause 720]
734. m1_subset_1(k1_xboole_0,k5_numbers) <- (11, 14) [forward demodulation 511,496] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
735. m1_subset_1(k6_numbers,k5_numbers) <- (11, 12, 14) [forward demodulation 734,501] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
737. 59 <=> m1_subset_1(k6_numbers,k5_numbers) [avatar definition]
739. m1_subset_1(k6_numbers,k5_numbers) <- (59) [avatar component clause 737] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
740. 59 | ~11 | ~12 | ~14 [avatar split clause 735,509,499,494,737]
747. v2_membered(k1_numbers) <- (7) [resolution 319,476] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
749. 61 <=> v2_membered(k1_numbers) [avatar definition]
751. v2_membered(k1_numbers) <- (61) [avatar component clause 749] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
752. 61 | ~7 [avatar split clause 747,474,749]
1005. 99 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1006. v1_xxreal_0(k6_numbers) <- (99) [avatar component clause 1005] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1149. v1_xxreal_0(k6_numbers) | ~v2_membered(k5_numbers) <- (59) [resolution 316,739] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1156. 123 <=> v2_membered(k5_numbers) [avatar definition]
1159. ~123 | 99 | ~59 [avatar split clause 1149,737,1005,1156]
1260. m1_subset_1(k6_numbers,k7_numbers) | ~m1_subset_1(sK0,k7_numbers) <- (2) [superposition 350,449] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1262. m1_subset_1(k6_numbers,k7_numbers) <- (2, 3) [subsumption resolution 1260,456] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1264. 138 <=> m1_subset_1(k6_numbers,k7_numbers) [avatar definition]
1267. 138 | ~2 | ~3 [avatar split clause 1262,454,448,1264]
1331. v2_membered(k5_numbers) | ~v2_membered(k1_numbers) <- (16) [resolution 317,521] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1333. v2_membered(k5_numbers) <- (16, 61) [subsumption resolution 1331,751] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1334. 123 | ~16 | ~61 [avatar split clause 1333,749,519,1156]
1425. m1_subset_1(k6_numbers,k7_numbers) <- (1, 3) [backward demodulation 456,445] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1427. k6_numbers != k3_extreal1(k6_numbers) <- (1, ~2) [forward demodulation 450,445] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1429. 144 <=> k6_numbers = k3_extreal1(k6_numbers) [avatar definition]
1431. k6_numbers != k3_extreal1(k6_numbers) <- (~144) [avatar component clause 1429] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1432. ~144 | ~1 | 2 [avatar split clause 1427,448,444,1429]
1443. 138 | ~1 | ~3 [avatar split clause 1425,454,444,1264]
1541. k6_numbers != k6_numbers | ~r1_xxreal_0(k6_numbers,k6_numbers) | ~m1_subset_1(k6_numbers,k7_numbers) <- (~144) [superposition 1431,354] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1542. ~r1_xxreal_0(k6_numbers,k6_numbers) | ~m1_subset_1(k6_numbers,k7_numbers) <- (~144) [trivial inequality removal 1541] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1547. 153 <=> r1_xxreal_0(k6_numbers,k6_numbers) [avatar definition]
1549. ~r1_xxreal_0(k6_numbers,k6_numbers) <- (~153) [avatar component clause 1547] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1550. ~138 | ~153 | 144 [avatar split clause 1542,1429,1547,1264]
1573. ~v1_xxreal_0(k6_numbers) <- (22, ~153) [resolution 1549,548] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1576. $false <- (22, 99, ~153) [subsumption resolution 1573,1006] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1577. ~22 | ~99 | 153 [avatar contradiction clause 1576]
1597. ~r1_xxreal_0(k6_numbers,k6_numbers) | k6_numbers = sK0 | ~m1_subset_1(sK0,k7_numbers) <- (2) [superposition 356,449] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1609. ~r1_xxreal_0(k6_numbers,k6_numbers) | ~m1_subset_1(sK0,k7_numbers) <- (~1, 2) [subsumption resolution 1597,446] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1610. ~r1_xxreal_0(k6_numbers,k6_numbers) <- (~1, 2, 3) [subsumption resolution 1609,456] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1611. ~153 | 1 | ~2 | ~3 [avatar split clause 1610,454,448,444,1547]
1617. $false [avatar sat refutation 451,452,457,477,497,502,512,522,549,569,721,740,752,1159,1267,1334,1432,1443,1550,1577,1611] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:874,thDist:-874}
% SZS output end Proof for extreal2__t5_extreal2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 122
% >>> Preprocessing
% Introduced skolems: 18
% Pure predicates: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 154
% Generated clauses: 707
% Active clauses: 185
% Passive clauses: 339
% Final active clauses: 163
% Final passive clauses: 125
% Main loop iterations started: 206
% 
% >>> Simplifying Inferences
% Duplicate literals: 10
% Trivial inequalities: 1
% Fw subsumption resolutions: 59
% Fw demodulations: 8
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 20
% Forward subsumptions: 108
% 
% >>> Generating Inferences
% Binary resolution: 221
% Unit resulting resolution: 3
% Forward superposition: 30
% Backward superposition: 33
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 38
% Split components: 81
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 183
% SAT solver unit clauses: 58
% SAT solver binary clauses: 35
% 
% Memory used [KB]: 5884
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% congruence closure: 0.002 s
% other: 0.005 s ( own 0.002 s ) 

