i: [1,95,0,2,1,1,0,0] t62_yellow_5
a: [3,154,0,2,1,95]
s: 154
i: [1,94,0,2,1,1,0,0] t62_yellow_5
a: [3,159,0,2,1,94]
s: 159
i: [1,93,0,2,1,1,0,0] t62_yellow_5
a: [3,164,0,2,1,93]
s: 164
i: [1,96,0,7,1,1,0,0] t62_yellow_5
a: [3,149,0,7,1,96]
s: 149
i: [1,144,0,2,1,0,0,1] existence_l1_orders_2
a: [3,169,0,2,1,144]
s: 169
i: [1,134,0,4,1,0,0,2] existence_m1_subset_1
s: 134
i: [1,99,0,6,3,0,0,255] cc4_yellow_0
s: 99
i: [1,100,0,6,3,0,0,255] cc4_yellow_0
s: 100
i: [1,101,0,6,3,0,0,255] cc1_lattice3
s: 101
i: [1,102,0,6,3,0,0,255] cc2_lattice3
s: 102
i: [1,97,0,7,2,0,0,1] dt_k4_yellow_0
s: 97
i: [1,98,0,7,2,0,0,1] dt_k3_yellow_0
s: 98
i: [1,104,0,8,4,0,0,255] cc6_waybel_1
s: 104
i: [1,106,0,8,4,0,0,255] cc7_waybel_1
s: 106
i: [1,108,0,8,4,0,0,255] cc5_waybel_1
s: 108
i: [1,109,0,8,4,0,0,255] cc5_waybel_1
s: 109
i: [1,110,0,8,4,0,0,255] cc5_waybel_1
s: 110
i: [1,111,0,8,4,0,0,255] cc5_waybel_1
s: 111
i: [1,112,0,8,4,0,0,255] cc5_waybel_1
s: 112
i: [1,116,0,8,4,0,0,255] cc10_waybel_1
s: 116
i: [1,118,0,8,4,0,0,255] cc8_waybel_1
s: 118
d: [2,171,1,6,3,1,40,116,154]
d: [2,173,1,4,2,2,21,171,159]
d: [2,174,1,2,1,3,21,173,164]
a: [3,178,1,2,1,174]
s: 178
i: [1,119,0,8,4,0,0,255] cc8_waybel_1
s: 119
d: [2,193,1,6,3,1,40,118,154]
d: [2,195,1,4,2,2,21,193,159]
d: [2,196,1,2,1,3,21,195,164]
a: [3,200,1,2,1,196]
s: 200
i: [1,120,0,8,4,0,0,255] cc8_waybel_1
s: 120
i: [1,121,0,8,4,0,0,255] cc8_waybel_1
s: 121
d: [2,207,2,6,3,1,40,178,104]
d: [2,214,2,4,2,2,21,207,164]
d: [2,215,2,2,1,3,21,214,154]
a: [3,219,2,2,1,215]
s: 219
i: [1,122,0,8,4,0,0,255] cc8_waybel_1
s: 122
d: [2,208,2,6,3,1,40,178,106]
d: [2,221,2,4,2,2,21,208,164]
d: [2,222,2,2,1,3,21,221,154]
a: [3,226,2,2,1,222]
s: 226
i: [1,123,0,8,4,0,0,255] cc8_waybel_1
s: 123
i: [1,124,0,8,4,0,0,255] cc8_waybel_1
s: 124
d: [2,299,1,6,3,1,40,123,154]
d: [2,301,1,4,2,2,21,299,159]
d: [2,302,1,2,1,3,21,301,164]
a: [3,306,1,2,1,302]
s: 306
i: [1,125,0,8,4,0,0,255] cc8_waybel_1
s: 125
d: [2,328,1,6,3,1,40,125,154]
d: [2,330,1,4,2,2,21,328,159]
d: [2,331,1,2,1,3,21,330,164]
a: [3,335,1,2,1,331]
s: 335
i: [1,126,0,8,4,0,0,255] cc5_yellow_0
s: 126
i: [1,135,0,14,4,0,0,1] dt_k7_waybel_1
s: 135
d: [2,210,2,6,3,1,40,178,109]
d: [2,231,2,4,2,2,21,210,164]
d: [2,232,2,2,1,3,21,231,154]
a: [3,236,2,2,1,232]
s: 236
i: [1,142,0,17,4,0,0,255] dt_k10_lattice3
s: 142
d: [2,211,2,6,3,1,40,178,110]
d: [2,238,2,4,2,2,21,211,164]
d: [2,239,2,2,1,3,21,238,154]
a: [3,243,2,2,1,239]
s: 243
i: [1,143,0,17,4,0,0,255] dt_k11_lattice3
s: 143
i: [1,136,0,21,6,0,0,255] dt_k12_lattice3
s: 136
d: [2,212,2,6,3,1,40,178,111]
d: [2,245,2,4,2,2,21,212,164]
d: [2,246,2,2,1,3,21,245,154]
a: [3,250,2,2,1,246]
s: 250
i: [1,137,0,21,6,0,0,255] dt_k13_lattice3
s: 137
d: [2,213,2,6,3,1,40,178,112]
d: [2,252,2,4,2,2,21,213,164]
d: [2,253,2,2,1,3,21,252,154]
a: [3,257,2,2,1,253]
s: 257
i: [1,114,0,22,11,0,0,255] cc9_waybel_1
s: 114
d: [2,345,0,20,10,0,21,114,101]
s: 345
d: [2,321,2,4,2,1,40,306,99]
d: [2,322,2,2,1,2,21,321,154]
a: [3,326,2,2,1,322]
s: 326
i: [1,127,0,22,7,0,0,255] t3_waybel_1
s: 127
d: [2,172,1,6,3,1,40,116,169]
a: [3,191,1,2,1,172]
s: 191
d: [2,347,0,20,6,0,21,127,102]
s: 347
i: [1,129,0,23,6,0,0,255] d23_waybel_1
s: 129
i: [1,130,0,23,6,0,0,255] d23_waybel_1
s: 130
i: [1,138,0,23,6,0,0,255] commutativity_k12_lattice3
s: 138
i: [1,139,0,23,6,0,0,255] commutativity_k13_lattice3
s: 139
i: [1,140,0,23,6,0,0,255] redefinition_k13_lattice3
s: 140
i: [1,141,0,23,6,0,0,255] redefinition_k12_lattice3
s: 141
d: [2,370,1,35,11,0,44,141,347]
d: [2,379,1,25,7,0,18,370]
d: [2,380,1,20,6,0,21,379,98]
s: 380
i: [1,128,0,25,9,0,0,255] t3_waybel_1
s: 128
d: [2,374,1,35,11,0,44,347,141]
d: [2,375,1,25,7,0,18,374]
d: [2,383,1,20,6,0,21,375,98]
s: 383
d: [2,398,0,23,8,0,21,128,101]
s: 398
i: [1,131,0,30,7,0,0,255] d23_waybel_1
s: 131
d: [2,384,2,35,11,0,44,380,138]
d: [2,393,2,25,7,0,18,384]
d: [2,394,2,20,6,0,21,393,98]
s: 394
i: [1,145,0,30,10,0,0,255] t11_yellow_5
s: 145
d: [2,385,2,35,11,0,44,380,138]
d: [2,392,2,25,7,0,18,385]
d: [2,395,2,20,6,0,21,392,98]
s: 395
i: [1,132,0,31,11,0,0,255] t11_yellow_5
s: 132
d: [2,438,1,32,11,1,44,149,132]
d: [2,444,1,30,10,2,21,438,154]
d: [2,445,1,28,9,3,21,444,159]
d: [2,446,1,26,8,4,21,445,257]
d: [2,447,1,24,7,5,21,446,250]
d: [2,448,1,22,6,6,21,447,243]
d: [2,449,1,20,5,7,21,448,236]
d: [2,450,1,18,4,8,21,449,200]
a: [3,454,1,5,1,450]
s: 454
d: [2,459,2,0,0,2,52,154,454,98]
e: 459
d: [2,460,2,2,1,1,40,454,98]
d: [2,462,2,0,0,2,21,460,154]
e: 462
r: 454
a: [3,453,1,5,1,450]
s: 453
a: [3,457,1,13,3,450]
s: 457
d: [2,464,2,11,2,1,47,457]
a: [3,468,2,6,1,464]
s: 468
d: [2,399,1,38,13,0,44,398,140]
d: [2,408,1,28,9,0,18,399]
d: [2,409,1,23,8,0,21,408,98]
s: 409
d: [2,401,1,38,13,0,44,140,398]
d: [2,406,1,28,9,0,18,401]
d: [2,412,1,23,8,0,21,406,98]
s: 412
d: [2,386,2,35,11,0,44,138,380]
d: [2,391,2,25,7,0,18,386]
d: [2,396,2,20,6,0,21,391,98]
s: 396
d: [2,434,1,32,12,0,40,145,135]
d: [2,435,1,26,10,0,18,434]
d: [2,436,1,24,9,0,21,435,101]
s: 436
r: 145
d: [2,387,2,35,11,0,44,138,380]
d: [2,390,2,25,7,0,18,387]
d: [2,397,2,20,6,0,21,390,98]
s: 397
d: [2,362,1,37,11,0,44,140,129]
d: [2,367,1,27,8,0,18,362]
d: [2,368,1,25,7,0,21,367,101]
s: 368
d: [2,363,1,37,11,0,44,129,140]
d: [2,366,1,27,8,0,18,363]
d: [2,369,1,25,7,0,21,366,101]
s: 369
d: [2,474,2,38,13,0,44,409,139]
d: [2,483,2,28,9,0,18,474]
d: [2,484,2,23,8,0,21,483,98]
s: 484
d: [2,371,1,37,11,0,44,141,130]
d: [2,378,1,27,8,0,18,371]
d: [2,381,1,25,7,0,21,378,102]
s: 381
d: [2,475,2,38,13,0,44,409,139]
d: [2,482,2,28,9,0,18,475]
d: [2,485,2,23,8,0,21,482,98]
s: 485
d: [2,372,1,37,11,0,44,130,141]
d: [2,377,1,27,8,0,18,372]
d: [2,382,1,25,7,0,21,377,102]
s: 382
d: [2,400,1,38,13,0,44,398,129]
d: [2,407,1,32,11,0,18,400]
d: [2,410,1,30,10,0,21,407,101]
d: [2,411,1,25,9,0,21,410,98]
s: 411
d: [2,476,2,38,13,0,44,139,409]
d: [2,481,2,28,9,0,18,476]
d: [2,486,2,23,8,0,21,481,98]
s: 486
d: [2,402,1,38,13,0,44,129,398]
d: [2,405,1,32,11,0,18,402]
d: [2,413,1,30,10,0,21,405,101]
d: [2,414,1,25,9,0,21,413,98]
s: 414
d: [2,477,2,38,13,0,44,139,409]
d: [2,480,2,28,9,0,18,477]
d: [2,487,2,23,8,0,21,480,98]
s: 487
d: [2,415,1,44,12,0,44,131,140]
d: [2,420,1,34,9,0,18,415]
d: [2,421,1,32,8,0,21,420,101]
s: 421
d: [2,417,1,46,14,0,44,131,398]
d: [2,418,1,40,12,0,18,417]
d: [2,422,1,38,11,0,21,418,101]
d: [2,423,1,33,10,0,21,422,98]
s: 423
d: [2,494,2,39,12,0,44,139,368]
d: [2,499,2,25,7,0,18,494]
s: 499
d: [2,437,1,55,21,0,44,132,132]
d: [2,443,1,37,13,0,18,437]
s: 443
d: [2,493,2,39,12,0,44,139,368]
d: [2,500,2,25,7,0,18,493]
s: 500
d: [2,491,2,39,12,0,44,368,139]
d: [2,502,2,25,7,0,18,491]
s: 502
d: [2,490,2,39,12,0,44,368,139]
d: [2,503,2,25,7,0,18,490]
s: 503
d: [2,492,2,40,14,0,44,368,409]
d: [2,501,2,30,10,0,18,492]
d: [2,504,2,25,9,0,21,501,98]
s: 504
d: [2,496,2,40,14,0,44,409,368]
d: [2,497,2,30,10,0,18,496]
d: [2,505,2,25,9,0,21,497,98]
s: 505
d: [2,525,2,39,12,0,44,138,381]
d: [2,528,2,25,7,0,18,525]
s: 528
d: [2,524,2,39,12,0,44,138,381]
d: [2,529,2,25,7,0,18,524]
s: 529
d: [2,523,2,39,12,0,44,381,138]
d: [2,530,2,25,7,0,18,523]
s: 530
d: [2,522,2,39,12,0,44,381,138]
d: [2,531,2,25,7,0,18,522]
s: 531
d: [2,547,2,42,15,0,44,423,347]
d: [2,552,2,32,11,0,18,547]
d: [2,553,2,27,10,0,19,552]
s: 553
d: [2,591,3,25,9,0,47,553]
d: [2,592,3,20,8,0,21,591,97]
s: 592
d: [2,593,4,0,0,8,52,257,243,326,250,200,236,154,468,592]
e: 593
d: [2,594,4,14,7,1,40,592,468]
d: [2,600,4,12,6,2,21,594,154]
d: [2,601,4,10,5,3,21,600,236]
d: [2,602,4,8,4,4,21,601,200]
d: [2,603,4,6,3,5,21,602,250]
d: [2,604,4,4,2,6,21,603,326]
d: [2,605,4,2,1,7,21,604,243]
d: [2,606,4,0,0,8,21,605,257]
e: 606
r: 468
a: [3,472,2,5,1,464]
s: 472
d: [2,608,3,0,0,2,52,154,472,97]
e: 608
d: [2,609,3,2,1,1,40,472,97]
d: [2,611,3,0,0,2,21,609,154]
e: 611
d: [2,259,1,6,3,1,40,119,154]
d: [2,261,1,4,2,2,21,259,159]
d: [2,262,1,2,1,3,21,261,164]
d: [2,269,1,6,3,1,40,120,154]
d: [2,271,1,4,2,2,21,269,159]
d: [2,272,1,2,1,3,21,271,164]
d: [2,279,1,6,3,1,40,121,154]
d: [2,281,1,4,2,2,21,279,159]
d: [2,282,1,2,1,3,21,281,164]
d: [2,289,1,6,3,1,40,122,154]
d: [2,291,1,4,2,2,21,289,159]
d: [2,292,1,2,1,3,21,291,164]
f: 611,606,464,462,450,322,302,292,282,272,262,196,93,94,95,96
% Refutation found. Thanks to Tanya!
% SZS status Theorem for yellow_5__t62_yellow_5
% SZS output start Proof for yellow_5__t62_yellow_5
1. ! [X0] : ((l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0)) => k7_waybel_1(X0,k3_yellow_0(X0)) = k4_yellow_0(X0)) [input t62_yellow_5]
2. ~! [X0] : ((l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0)) => k7_waybel_1(X0,k3_yellow_0(X0)) = k4_yellow_0(X0)) [negated conjecture 1]
4. ! [X0] : (l1_orders_2(X0) => (v1_lattice3(X0) => ~v2_struct_0(X0))) [input cc1_lattice3]
5. ! [X0] : (l1_orders_2(X0) => (v2_lattice3(X0) => ~v2_struct_0(X0))) [input cc2_lattice3]
6. ! [X0] : (l1_orders_2(X0) => (v3_yellow_0(X0) => (v2_yellow_0(X0) & v1_yellow_0(X0)))) [input cc4_yellow_0]
11. ! [X0] : (l1_orders_2(X0) => ((v11_waybel_1(X0) & ~v2_struct_0(X0)) => (v10_waybel_1(X0) & v2_waybel_1(X0) & v3_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & ~v2_struct_0(X0)))) [input cc8_waybel_1]
15. ! [X0] : ((l1_orders_2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r6_waybel_1(X0,X1,X2) <=> (k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2)))))) [input d23_waybel_1]
20. ! [X0] : (l1_orders_2(X0) => m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0))) [input dt_k3_yellow_0]
21. ! [X0] : (l1_orders_2(X0) => m1_subset_1(k4_yellow_0(X0),u1_struct_0(X0))) [input dt_k4_yellow_0]
32. ! [X0] : ((l1_orders_2(X0) & v11_waybel_1(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r6_waybel_1(X0,X1,X2) <=> k7_waybel_1(X0,X1) = X2)))) [input t11_yellow_5]
33. ! [X0] : ((l1_orders_2(X0) & v1_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => (((v4_orders_2(X0) & v3_orders_2(X0) & v1_lattice3(X0)) => k10_lattice3(X0,k3_yellow_0(X0),X1) = X1) & (v2_lattice3(X0) => k3_yellow_0(X0) = k11_lattice3(X0,k3_yellow_0(X0),X1))))) [input t3_waybel_1]
36. ? [X0] : (k7_waybel_1(X0,k3_yellow_0(X0)) != k4_yellow_0(X0) & (l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
37. ? [X0] : (k7_waybel_1(X0,k3_yellow_0(X0)) != k4_yellow_0(X0) & l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0)) [flattening 36]
38. ! [X0] : (m1_subset_1(k4_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0)) [ennf transformation 21]
39. ! [X0] : (m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0)) [ennf transformation 20]
40. ! [X0] : (((v2_yellow_0(X0) & v1_yellow_0(X0)) | ~v3_yellow_0(X0)) | ~l1_orders_2(X0)) [ennf transformation 6]
41. ! [X0] : ((v2_yellow_0(X0) & v1_yellow_0(X0)) | ~v3_yellow_0(X0) | ~l1_orders_2(X0)) [flattening 40]
42. ! [X0] : ((~v2_struct_0(X0) | ~v1_lattice3(X0)) | ~l1_orders_2(X0)) [ennf transformation 4]
43. ! [X0] : (~v2_struct_0(X0) | ~v1_lattice3(X0) | ~l1_orders_2(X0)) [flattening 42]
44. ! [X0] : ((~v2_struct_0(X0) | ~v2_lattice3(X0)) | ~l1_orders_2(X0)) [ennf transformation 5]
45. ! [X0] : (~v2_struct_0(X0) | ~v2_lattice3(X0) | ~l1_orders_2(X0)) [flattening 44]
56. ! [X0] : (((v10_waybel_1(X0) & v2_waybel_1(X0) & v3_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & ~v2_struct_0(X0)) | (~v11_waybel_1(X0) | v2_struct_0(X0))) | ~l1_orders_2(X0)) [ennf transformation 11]
57. ! [X0] : ((v10_waybel_1(X0) & v2_waybel_1(X0) & v3_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & ~v2_struct_0(X0)) | ~v11_waybel_1(X0) | v2_struct_0(X0) | ~l1_orders_2(X0)) [flattening 56]
60. ! [X0] : (! [X1] : (((k10_lattice3(X0,k3_yellow_0(X0),X1) = X1 | (~v4_orders_2(X0) | ~v3_orders_2(X0) | ~v1_lattice3(X0))) & (k3_yellow_0(X0) = k11_lattice3(X0,k3_yellow_0(X0),X1) | ~v2_lattice3(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 33]
61. ! [X0] : (! [X1] : (((k10_lattice3(X0,k3_yellow_0(X0),X1) = X1 | ~v4_orders_2(X0) | ~v3_orders_2(X0) | ~v1_lattice3(X0)) & (k3_yellow_0(X0) = k11_lattice3(X0,k3_yellow_0(X0),X1) | ~v2_lattice3(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0)) [flattening 60]
62. ! [X0] : (! [X1] : (! [X2] : ((r6_waybel_1(X0,X1,X2) <=> (k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 15]
63. ! [X0] : (! [X1] : (! [X2] : ((r6_waybel_1(X0,X1,X2) <=> (k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [flattening 62]
64. ! [X0] : (! [X1] : (! [X2] : ((r6_waybel_1(X0,X1,X2) <=> k7_waybel_1(X0,X1) = X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0))) [ennf transformation 32]
65. ! [X0] : (! [X1] : (! [X2] : ((r6_waybel_1(X0,X1,X2) <=> k7_waybel_1(X0,X1) = X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0)) [flattening 64]
84. ? [X0] : (k7_waybel_1(X0,k3_yellow_0(X0)) != k4_yellow_0(X0) & l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0)) => (k7_waybel_1(sK0,k3_yellow_0(sK0)) != k4_yellow_0(sK0) & l1_orders_2(sK0) & v11_waybel_1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
85. k7_waybel_1(sK0,k3_yellow_0(sK0)) != k4_yellow_0(sK0) & l1_orders_2(sK0) & v11_waybel_1(sK0) & ~v2_struct_0(sK0) [skolemisation 37,84]
86. ! [X0] : (! [X1] : (! [X2] : (((r6_waybel_1(X0,X1,X2) | (k3_yellow_0(X0) != k11_lattice3(X0,X1,X2) | k4_yellow_0(X0) != k10_lattice3(X0,X1,X2))) & ((k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2)) | ~r6_waybel_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [nnf transformation 63]
87. ! [X0] : (! [X1] : (! [X2] : (((r6_waybel_1(X0,X1,X2) | k3_yellow_0(X0) != k11_lattice3(X0,X1,X2) | k4_yellow_0(X0) != k10_lattice3(X0,X1,X2)) & ((k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2)) | ~r6_waybel_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [flattening 86]
88. ! [X0] : (! [X1] : (! [X2] : (((r6_waybel_1(X0,X1,X2) | k7_waybel_1(X0,X1) != X2) & (k7_waybel_1(X0,X1) = X2 | ~r6_waybel_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0)) [nnf transformation 65]
93. ~v2_struct_0(sK0) [cnf transformation 85] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
94. v11_waybel_1(sK0) [cnf transformation 85] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
95. l1_orders_2(sK0) [cnf transformation 85] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
96. k7_waybel_1(sK0,k3_yellow_0(sK0)) != k4_yellow_0(sK0) [cnf transformation 85] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
97. m1_subset_1(k4_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 38] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
98. m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 39] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
99. ~v3_yellow_0(X0) | v1_yellow_0(X0) | ~l1_orders_2(X0) [cnf transformation 41] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
101. ~v1_lattice3(X0) | ~v2_struct_0(X0) | ~l1_orders_2(X0) [cnf transformation 43] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
102. ~v2_lattice3(X0) | ~v2_struct_0(X0) | ~l1_orders_2(X0) [cnf transformation 45] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
118. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v3_orders_2(X0) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
119. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v4_orders_2(X0) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
120. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v5_orders_2(X0) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
121. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v1_lattice3(X0) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
122. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v2_lattice3(X0) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
123. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v3_yellow_0(X0) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
127. k3_yellow_0(X0) = k11_lattice3(X0,k3_yellow_0(X0),X1) | ~v2_lattice3(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0) [cnf transformation 61] {a:0,w:22,wCS:8800,sine:255,thAx:0,allAx:1,thDist:-1}
128. k10_lattice3(X0,k3_yellow_0(X0),X1) = X1 | ~v4_orders_2(X0) | ~v3_orders_2(X0) | ~v1_lattice3(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0) [cnf transformation 61] {a:0,w:25,wCS:10000,sine:255,thAx:0,allAx:1,thDist:-1}
131. k4_yellow_0(X0) != k10_lattice3(X0,X1,X2) | k3_yellow_0(X0) != k11_lattice3(X0,X1,X2) | r6_waybel_1(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0) [cnf transformation 87] {a:0,w:30,wCS:12000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
132. k7_waybel_1(X0,X1) = X2 | ~r6_waybel_1(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) [cnf transformation 88] {a:0,w:31,wCS:12400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
147. 1 <=> k7_waybel_1(sK0,k3_yellow_0(sK0)) = k4_yellow_0(sK0) [avatar definition]
149. k7_waybel_1(sK0,k3_yellow_0(sK0)) != k4_yellow_0(sK0) <- (~1) [avatar component clause 147] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
150. ~1 [avatar split clause 96,147]
152. 2 <=> l1_orders_2(sK0) [avatar definition]
154. l1_orders_2(sK0) <- (2) [avatar component clause 152] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
155. 2 [avatar split clause 95,152]
157. 3 <=> v11_waybel_1(sK0) [avatar definition]
159. v11_waybel_1(sK0) <- (3) [avatar component clause 157] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
160. 3 [avatar split clause 94,157]
162. 4 <=> v2_struct_0(sK0) [avatar definition]
164. ~v2_struct_0(sK0) <- (~4) [avatar component clause 162] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
165. ~4 [avatar split clause 93,162]
193. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v3_orders_2(sK0) <- (2) [resolution 118,154] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
195. v2_struct_0(sK0) | v3_orders_2(sK0) <- (2, 3) [subsumption resolution 193,159] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
196. v3_orders_2(sK0) <- (2, 3, ~4) [subsumption resolution 195,164] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
198. 10 <=> v3_orders_2(sK0) [avatar definition]
200. v3_orders_2(sK0) <- (10) [avatar component clause 198] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
201. 10 | ~2 | ~3 | 4 [avatar split clause 196,162,157,152,198]
234. 14 <=> v4_orders_2(sK0) [avatar definition]
236. v4_orders_2(sK0) <- (14) [avatar component clause 234] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
241. 15 <=> v5_orders_2(sK0) [avatar definition]
243. v5_orders_2(sK0) <- (15) [avatar component clause 241] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
248. 16 <=> v1_lattice3(sK0) [avatar definition]
250. v1_lattice3(sK0) <- (16) [avatar component clause 248] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
255. 17 <=> v2_lattice3(sK0) [avatar definition]
257. v2_lattice3(sK0) <- (17) [avatar component clause 255] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
259. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v4_orders_2(sK0) <- (2) [resolution 119,154] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
261. v2_struct_0(sK0) | v4_orders_2(sK0) <- (2, 3) [subsumption resolution 259,159] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
262. v4_orders_2(sK0) <- (2, 3, ~4) [subsumption resolution 261,164] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
263. 14 | ~2 | ~3 | 4 [avatar split clause 262,162,157,152,234]
269. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v5_orders_2(sK0) <- (2) [resolution 120,154] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
271. v2_struct_0(sK0) | v5_orders_2(sK0) <- (2, 3) [subsumption resolution 269,159] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
272. v5_orders_2(sK0) <- (2, 3, ~4) [subsumption resolution 271,164] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
273. 15 | ~2 | ~3 | 4 [avatar split clause 272,162,157,152,241]
279. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v1_lattice3(sK0) <- (2) [resolution 121,154] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
281. v2_struct_0(sK0) | v1_lattice3(sK0) <- (2, 3) [subsumption resolution 279,159] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
282. v1_lattice3(sK0) <- (2, 3, ~4) [subsumption resolution 281,164] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
283. 16 | ~2 | ~3 | 4 [avatar split clause 282,162,157,152,248]
289. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v2_lattice3(sK0) <- (2) [resolution 122,154] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
291. v2_struct_0(sK0) | v2_lattice3(sK0) <- (2, 3) [subsumption resolution 289,159] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
292. v2_lattice3(sK0) <- (2, 3, ~4) [subsumption resolution 291,164] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
293. 17 | ~2 | ~3 | 4 [avatar split clause 292,162,157,152,255]
299. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v3_yellow_0(sK0) <- (2) [resolution 123,154] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
301. v2_struct_0(sK0) | v3_yellow_0(sK0) <- (2, 3) [subsumption resolution 299,159] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
302. v3_yellow_0(sK0) <- (2, 3, ~4) [subsumption resolution 301,164] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
304. 22 <=> v3_yellow_0(sK0) [avatar definition]
306. v3_yellow_0(sK0) <- (22) [avatar component clause 304] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
307. 22 | ~2 | ~3 | 4 [avatar split clause 302,162,157,152,304]
321. v1_yellow_0(sK0) | ~l1_orders_2(sK0) <- (22) [resolution 306,99] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
322. v1_yellow_0(sK0) <- (2, 22) [subsumption resolution 321,154] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
324. 25 <=> v1_yellow_0(sK0) [avatar definition]
326. v1_yellow_0(sK0) <- (25) [avatar component clause 324] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
327. 25 | ~2 | ~22 [avatar split clause 322,304,152,324]
347. k3_yellow_0(X0) = k11_lattice3(X0,k3_yellow_0(X0),X1) | ~v2_lattice3(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) [subsumption resolution 127,102] {a:0,w:20,wCS:8000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
398. k10_lattice3(X0,k3_yellow_0(X0),X1) = X1 | ~v4_orders_2(X0) | ~v3_orders_2(X0) | ~v1_lattice3(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) [subsumption resolution 128,101] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
417. k4_yellow_0(X6) != X7 | k3_yellow_0(X6) != k11_lattice3(X6,k3_yellow_0(X6),X7) | r6_waybel_1(X6,k3_yellow_0(X6),X7) | ~m1_subset_1(X7,u1_struct_0(X6)) | ~m1_subset_1(k3_yellow_0(X6),u1_struct_0(X6)) | ~l1_orders_2(X6) | v2_struct_0(X6) | ~v4_orders_2(X6) | ~v3_orders_2(X6) | ~v1_lattice3(X6) | ~m1_subset_1(X7,u1_struct_0(X6)) | ~l1_orders_2(X6) | ~v1_yellow_0(X6) | ~v5_orders_2(X6) [superposition 131,398] {a:1,w:46,wCS:18400,sine:255,thAx:0,allAx:2,thDist:-2}
418. k4_yellow_0(X6) != X7 | k3_yellow_0(X6) != k11_lattice3(X6,k3_yellow_0(X6),X7) | r6_waybel_1(X6,k3_yellow_0(X6),X7) | ~m1_subset_1(X7,u1_struct_0(X6)) | ~m1_subset_1(k3_yellow_0(X6),u1_struct_0(X6)) | ~l1_orders_2(X6) | v2_struct_0(X6) | ~v4_orders_2(X6) | ~v3_orders_2(X6) | ~v1_lattice3(X6) | ~v1_yellow_0(X6) | ~v5_orders_2(X6) [duplicate literal removal 417] {a:1,w:40,wCS:16000,sine:255,thAx:0,allAx:2,thDist:-2}
422. k4_yellow_0(X6) != X7 | k3_yellow_0(X6) != k11_lattice3(X6,k3_yellow_0(X6),X7) | r6_waybel_1(X6,k3_yellow_0(X6),X7) | ~m1_subset_1(X7,u1_struct_0(X6)) | ~m1_subset_1(k3_yellow_0(X6),u1_struct_0(X6)) | ~l1_orders_2(X6) | ~v4_orders_2(X6) | ~v3_orders_2(X6) | ~v1_lattice3(X6) | ~v1_yellow_0(X6) | ~v5_orders_2(X6) [subsumption resolution 418,101] {a:1,w:38,wCS:15200,sine:255,thAx:0,allAx:2,thDist:-2}
423. k3_yellow_0(X6) != k11_lattice3(X6,k3_yellow_0(X6),X7) | k4_yellow_0(X6) != X7 | r6_waybel_1(X6,k3_yellow_0(X6),X7) | ~m1_subset_1(X7,u1_struct_0(X6)) | ~l1_orders_2(X6) | ~v4_orders_2(X6) | ~v3_orders_2(X6) | ~v1_lattice3(X6) | ~v1_yellow_0(X6) | ~v5_orders_2(X6) [subsumption resolution 422,98] {a:1,w:33,wCS:13200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
438. k4_yellow_0(sK0) != X0 | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v11_waybel_1(sK0) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1) [superposition 149,132] {a:1,w:32,wCS:3200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
444. k4_yellow_0(sK0) != X0 | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) | ~v11_waybel_1(sK0) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 2) [subsumption resolution 438,154] {a:1,w:30,wCS:3000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
445. k4_yellow_0(sK0) != X0 | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 2, 3) [subsumption resolution 444,159] {a:1,w:28,wCS:2800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
446. k4_yellow_0(sK0) != X0 | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 2, 3, 17) [subsumption resolution 445,257] {a:1,w:26,wCS:2600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
447. k4_yellow_0(sK0) != X0 | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 2, 3, 16, 17) [subsumption resolution 446,250] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
448. k4_yellow_0(sK0) != X0 | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 2, 3, 15, 16, 17) [subsumption resolution 447,243] {a:1,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
449. k4_yellow_0(sK0) != X0 | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) | ~v3_orders_2(sK0) <- (~1, 2, 3, 14, 15, 16, 17) [subsumption resolution 448,236] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
450. k4_yellow_0(sK0) != X0 | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) <- (~1, 2, 3, 10, 14, 15, 16, 17) [subsumption resolution 449,200] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
452. 28 <=> m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) [avatar definition]
454. ~m1_subset_1(k3_yellow_0(sK0),u1_struct_0(sK0)) <- (~28) [avatar component clause 452] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
456. 29 <=> ! [X0] : (k4_yellow_0(sK0) != X0 | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0)) [avatar definition]
457. k4_yellow_0(sK0) != X0 | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~r6_waybel_1(sK0,k3_yellow_0(sK0),X0) <- (29) [avatar component clause 456] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
458. ~28 | 29 | 1 | ~2 | ~3 | ~10 | ~14 | ~15 | ~16 | ~17 [avatar split clause 450,255,248,241,234,198,157,152,147,456,452]
460. ~l1_orders_2(sK0) <- (~28) [resolution 454,98] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
462. $false <- (2, ~28) [subsumption resolution 460,154] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
463. ~2 | 28 [avatar contradiction clause 462]
464. ~m1_subset_1(k4_yellow_0(sK0),u1_struct_0(sK0)) | ~r6_waybel_1(sK0,k3_yellow_0(sK0),k4_yellow_0(sK0)) <- (29) [equality resolution 457] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
466. 30 <=> r6_waybel_1(sK0,k3_yellow_0(sK0),k4_yellow_0(sK0)) [avatar definition]
468. ~r6_waybel_1(sK0,k3_yellow_0(sK0),k4_yellow_0(sK0)) <- (~30) [avatar component clause 466] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
470. 31 <=> m1_subset_1(k4_yellow_0(sK0),u1_struct_0(sK0)) [avatar definition]
472. ~m1_subset_1(k4_yellow_0(sK0),u1_struct_0(sK0)) <- (~31) [avatar component clause 470] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
473. ~30 | ~31 | ~29 [avatar split clause 464,456,470,466]
547. k3_yellow_0(X0) != k3_yellow_0(X0) | k4_yellow_0(X0) != X1 | r6_waybel_1(X0,k3_yellow_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) | ~v1_lattice3(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | ~v2_lattice3(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) [superposition 423,347] {a:2,w:42,wCS:16800,sine:1,thAx:0,allAx:3,thDist:-3}
552. k3_yellow_0(X0) != k3_yellow_0(X0) | k4_yellow_0(X0) != X1 | r6_waybel_1(X0,k3_yellow_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) | ~v1_lattice3(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | ~v2_lattice3(X0) [duplicate literal removal 547] {a:2,w:32,wCS:12800,sine:1,thAx:0,allAx:3,thDist:-3}
553. k4_yellow_0(X0) != X1 | r6_waybel_1(X0,k3_yellow_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) | ~v1_lattice3(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | ~v2_lattice3(X0) [trivial inequality removal 552] {a:2,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
591. r6_waybel_1(X0,k3_yellow_0(X0),k4_yellow_0(X0)) | ~m1_subset_1(k4_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) | ~v1_lattice3(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | ~v2_lattice3(X0) [equality resolution 553] {a:3,w:25,wCS:10000,sine:1,thAx:0,allAx:3,thDist:-3}
592. r6_waybel_1(X0,k3_yellow_0(X0),k4_yellow_0(X0)) | ~l1_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) | ~v1_lattice3(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | ~v2_lattice3(X0) [subsumption resolution 591,97] {a:3,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
594. ~l1_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) | ~v1_lattice3(sK0) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | ~v2_lattice3(sK0) <- (~30) [resolution 592,468] {a:4,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
600. ~v4_orders_2(sK0) | ~v3_orders_2(sK0) | ~v1_lattice3(sK0) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | ~v2_lattice3(sK0) <- (2, ~30) [subsumption resolution 594,154] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
601. ~v3_orders_2(sK0) | ~v1_lattice3(sK0) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | ~v2_lattice3(sK0) <- (2, 14, ~30) [subsumption resolution 600,236] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
602. ~v1_lattice3(sK0) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | ~v2_lattice3(sK0) <- (2, 10, 14, ~30) [subsumption resolution 601,200] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
603. ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | ~v2_lattice3(sK0) <- (2, 10, 14, 16, ~30) [subsumption resolution 602,250] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
604. ~v5_orders_2(sK0) | ~v2_lattice3(sK0) <- (2, 10, 14, 16, 25, ~30) [subsumption resolution 603,326] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
605. ~v2_lattice3(sK0) <- (2, 10, 14, 15, 16, 25, ~30) [subsumption resolution 604,243] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
606. $false <- (2, 10, 14, 15, 16, 17, 25, ~30) [subsumption resolution 605,257] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
607. ~2 | ~10 | ~14 | ~15 | ~16 | ~17 | ~25 | 30 [avatar contradiction clause 606]
609. ~l1_orders_2(sK0) <- (~31) [resolution 472,97] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
611. $false <- (2, ~31) [subsumption resolution 609,154] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
612. ~2 | 31 [avatar contradiction clause 611]
613. $false [avatar sat refutation 150,155,160,165,201,263,273,283,293,307,327,458,463,473,607,612] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:273,thDist:-273}
% SZS output end Proof for yellow_5__t62_yellow_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 32
% >>> Preprocessing
% Introduced skolems: 3
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 52
% Generated clauses: 381
% Active clauses: 78
% Passive clauses: 112
% Final active clauses: 75
% Final passive clauses: 12
% Main loop iterations started: 100
% 
% >>> Simplifying Inferences
% Duplicate literals: 380
% Trivial inequalities: 2
% Fw subsumption resolutions: 86
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Equational tautologies: 1
% Forward subsumptions: 57
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 39
% Unit resulting resolution: 3
% Forward superposition: 37
% Backward superposition: 47
% Self superposition: 1
% Equality resolution: 2
% 
% >>> AVATAR
% Split clauses: 11
% Split components: 31
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 40
% SAT solver unit clauses: 5
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9722
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

