i: [1,47,0,2,1,1,0,0] t2_xtuple_0
a: [3,85,0,2,1,47]
s: 85
i: [1,46,0,2,1,1,0,0] t2_xtuple_0
a: [3,90,0,2,1,46]
s: 90
i: [1,50,0,3,1,1,0,0] t2_xtuple_0
a: [3,70,0,3,1,50]
s: 70
i: [1,49,0,5,1,1,0,0] t2_xtuple_0
a: [3,75,0,5,1,49]
s: 75
i: [1,48,0,5,1,1,0,0] t2_xtuple_0
a: [3,80,0,5,1,48]
s: 80
i: [1,51,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,95,0,2,1,51]
s: 95
i: [1,52,0,2,1,0,0,255] fc1_xboole_0
a: [3,100,0,2,1,52]
s: 100
i: [1,64,0,2,1,0,0,2] rc2_xboole_0
a: [3,110,0,2,1,64]
s: 110
i: [1,65,0,2,1,0,0,2] rc1_xboole_0
a: [3,115,0,2,1,65]
s: 115
i: [1,66,0,2,1,0,0,1] rc1_xtuple_0
a: [3,120,0,2,1,66]
s: 120
i: [1,54,0,3,1,0,0,255] fc2_xboole_0
s: 54
i: [1,53,0,3,1,0,0,255] d2_xboole_0
a: [3,105,0,3,1,53]
s: 105
r: 95
i: [1,59,0,4,1,0,0,255] fc3_xboole_0
s: 59
i: [1,60,0,4,1,0,0,1] fc1_xtuple_0
s: 60
i: [1,55,0,5,2,0,0,255] t6_boole
s: 55
i: [1,61,0,6,1,0,0,1] rd2_xtuple_0
s: 61
i: [1,62,0,6,1,0,0,1] rd1_xtuple_0
s: 62
i: [1,58,0,7,3,0,0,1] t8_boole
s: 58
i: [1,63,0,7,1,0,0,255] commutativity_k2_tarski
s: 63
i: [1,56,0,9,2,0,0,1] rd3_xtuple_0
s: 56
i: [1,57,0,10,1,0,0,255] d5_tarski
s: 57
d: [2,127,1,9,2,1,44,56,80]
d: [2,132,1,7,1,2,21,127,85]
d: [2,133,1,7,1,3,22,132,75]
a: [3,137,1,7,1,133]
s: 137
d: [2,152,2,5,2,1,44,137,56]
d: [2,157,2,2,1,2,21,152,70]
d: [2,158,2,0,0,3,21,157,90]
e: 158
d: [2,153,2,5,2,1,44,56,137]
d: [2,160,2,2,1,2,21,153,70]
d: [2,161,2,0,0,3,21,160,90]
e: 161
d: [2,129,1,9,2,1,44,56,75]
d: [2,140,1,7,1,2,21,129,85]
d: [2,141,1,7,1,3,22,140,80]
f: 161,141,46,47,48,49,50
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t2_xtuple_0
% SZS output start Proof for xtuple_0__t2_xtuple_0
1. ! [X0] : (v1_xtuple_0(X0) => ! [X1] : (v1_xtuple_0(X1) => ((k2_xtuple_0(X0) = k2_xtuple_0(X1) & k1_xtuple_0(X0) = k1_xtuple_0(X1)) => X0 = X1))) [input t2_xtuple_0]
2. ~! [X0] : (v1_xtuple_0(X0) => ! [X1] : (v1_xtuple_0(X1) => ((k2_xtuple_0(X0) = k2_xtuple_0(X1) & k1_xtuple_0(X0) = k1_xtuple_0(X1)) => X0 = X1))) [negated conjecture 1]
23. ! [X0] : (v1_xtuple_0(X0) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0) [input rd3_xtuple_0]
32. ? [X0] : (? [X1] : ((X0 != X1 & (k2_xtuple_0(X0) = k2_xtuple_0(X1) & k1_xtuple_0(X0) = k1_xtuple_0(X1))) & v1_xtuple_0(X1)) & v1_xtuple_0(X0)) [ennf transformation 2]
33. ? [X0] : (? [X1] : (X0 != X1 & k2_xtuple_0(X0) = k2_xtuple_0(X1) & k1_xtuple_0(X0) = k1_xtuple_0(X1) & v1_xtuple_0(X1)) & v1_xtuple_0(X0)) [flattening 32]
35. ! [X0] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0)) [ennf transformation 23]
37. ? [X0] : (? [X1] : (X0 != X1 & k2_xtuple_0(X0) = k2_xtuple_0(X1) & k1_xtuple_0(X0) = k1_xtuple_0(X1) & v1_xtuple_0(X1)) & v1_xtuple_0(X0)) => (? [X1] : (sK0 != X1 & k2_xtuple_0(X1) = k2_xtuple_0(sK0) & k1_xtuple_0(X1) = k1_xtuple_0(sK0) & v1_xtuple_0(X1)) & v1_xtuple_0(sK0)) [choice axiom]
38. ? [X1] : (sK0 != X1 & k2_xtuple_0(X1) = k2_xtuple_0(sK0) & k1_xtuple_0(X1) = k1_xtuple_0(sK0) & v1_xtuple_0(X1)) => (sK0 != sK1 & k2_xtuple_0(sK0) = k2_xtuple_0(sK1) & k1_xtuple_0(sK0) = k1_xtuple_0(sK1) & v1_xtuple_0(sK1)) [choice axiom]
39. (sK0 != sK1 & k2_xtuple_0(sK0) = k2_xtuple_0(sK1) & k1_xtuple_0(sK0) = k1_xtuple_0(sK1) & v1_xtuple_0(sK1)) & v1_xtuple_0(sK0) [skolemisation 33,38,37]
46. v1_xtuple_0(sK0) [cnf transformation 39] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
47. v1_xtuple_0(sK1) [cnf transformation 39] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. k1_xtuple_0(sK0) = k1_xtuple_0(sK1) [cnf transformation 39] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. k2_xtuple_0(sK0) = k2_xtuple_0(sK1) [cnf transformation 39] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. sK0 != sK1 [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0) [cnf transformation 35] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. 1 <=> sK0 = sK1 [avatar definition]
70. sK0 != sK1 <- (~1) [avatar component clause 68] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
71. ~1 [avatar split clause 50,68]
73. 2 <=> k2_xtuple_0(sK0) = k2_xtuple_0(sK1) [avatar definition]
75. k2_xtuple_0(sK0) = k2_xtuple_0(sK1) <- (2) [avatar component clause 73] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
76. 2 [avatar split clause 49,73]
78. 3 <=> k1_xtuple_0(sK0) = k1_xtuple_0(sK1) [avatar definition]
80. k1_xtuple_0(sK0) = k1_xtuple_0(sK1) <- (3) [avatar component clause 78] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
81. 3 [avatar split clause 48,78]
83. 4 <=> v1_xtuple_0(sK1) [avatar definition]
85. v1_xtuple_0(sK1) <- (4) [avatar component clause 83] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. 4 [avatar split clause 47,83]
88. 5 <=> v1_xtuple_0(sK0) [avatar definition]
90. v1_xtuple_0(sK0) <- (5) [avatar component clause 88] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. 5 [avatar split clause 46,88]
129. sK1 = k4_tarski(k1_xtuple_0(sK1),k2_xtuple_0(sK0)) | ~v1_xtuple_0(sK1) <- (2) [superposition 56,75] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
135. 12 <=> sK1 = k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) [avatar definition]
137. sK1 = k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (12) [avatar component clause 135] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
140. sK1 = k4_tarski(k1_xtuple_0(sK1),k2_xtuple_0(sK0)) <- (2, 4) [subsumption resolution 129,85] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
141. sK1 = k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (2, 3, 4) [forward demodulation 140,80] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
142. 12 | ~2 | ~3 | ~4 [avatar split clause 141,83,78,73,135]
153. sK0 = sK1 | ~v1_xtuple_0(sK0) <- (12) [superposition 56,137] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
160. ~v1_xtuple_0(sK0) <- (~1, 12) [subsumption resolution 153,70] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
161. $false <- (~1, 5, 12) [subsumption resolution 160,90] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
162. 1 | ~5 | ~12 [avatar contradiction clause 161]
163. $false [avatar sat refutation 71,76,81,86,91,142,162] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:61,thDist:-61}
% SZS output end Proof for xtuple_0__t2_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 66
% Active clauses: 22
% Passive clauses: 31
% Final active clauses: 21
% Final passive clauses: 9
% Main loop iterations started: 22
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 8
% Fw demodulations: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 6
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Forward superposition: 9
% Backward superposition: 11
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 11
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% binary resolution index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

