i: [1,19,0,2,1,1,0,0] t42_rewrite1
a: [3,57,0,2,1,19]
s: 57
i: [1,24,0,4,1,1,0,0] t42_rewrite1
a: [3,32,0,4,1,24]
s: 32
i: [1,23,0,8,2,1,0,0] t42_rewrite1
a: [3,41,0,4,1,23]
s: 41
i: [1,21,0,8,2,1,0,0] t42_rewrite1
a: [3,51,0,4,1,21]
s: 51
i: [1,25,0,13,3,0,0,1] d7_rewrite1
s: 25
i: [1,26,0,13,3,0,0,1] d7_rewrite1
s: 26
i: [1,27,0,14,4,0,0,1] d7_rewrite1
s: 27
i: [1,28,0,14,4,0,0,1] t16_rewrite1
s: 28
d: [2,59,1,10,3,1,40,27,41]
d: [2,64,1,8,2,2,21,59,57]
s: 64
d: [2,65,1,10,3,1,40,28,41]
d: [2,70,1,8,2,2,21,65,57]
s: 70
d: [2,61,1,19,5,0,40,27,26]
d: [2,62,1,17,4,0,18,61]
s: 62
d: [2,71,2,4,1,2,40,64,41]
a: [3,75,2,4,1,71]
s: 75
d: [2,60,1,19,5,0,40,27,25]
d: [2,63,1,17,4,0,18,60]
s: 63
d: [2,78,2,16,5,0,40,62,26]
d: [2,79,2,10,3,0,18,78]
s: 79
d: [2,67,1,19,5,0,40,28,26]
d: [2,68,1,17,4,0,18,67]
s: 68
d: [2,66,1,19,5,0,40,28,25]
d: [2,69,1,17,4,0,18,66]
s: 69
d: [2,85,3,6,2,1,40,79,51]
d: [2,87,3,4,1,2,21,85,57]
a: [3,91,3,4,1,87]
s: 91
d: [2,82,2,16,5,0,40,63,26]
d: [2,83,2,10,3,0,18,82]
s: 83
d: [2,110,3,6,2,1,40,83,51]
d: [2,113,3,4,1,2,21,110,57]
a: [3,117,3,4,1,113]
s: 117
d: [2,81,2,16,5,0,40,63,25]
d: [2,84,2,10,3,0,18,81]
s: 84
d: [2,94,2,20,6,0,40,68,62]
d: [2,99,2,14,4,0,18,94]
s: 99
d: [2,120,4,6,2,1,40,117,79]
d: [2,121,4,4,1,2,21,120,57]
a: [3,125,4,4,1,121]
s: 125
d: [2,93,2,20,6,0,40,68,63]
d: [2,100,2,14,4,0,18,93]
s: 100
d: [2,133,3,10,3,1,40,99,51]
d: [2,137,3,8,2,2,21,133,57]
s: 137
d: [2,102,2,20,6,0,40,69,62]
d: [2,107,2,14,4,0,18,102]
s: 107
d: [2,101,2,20,6,0,40,69,63]
d: [2,108,2,14,4,0,18,101]
s: 108
d: [2,156,4,4,1,3,40,137,41]
a: [3,160,4,4,1,156]
s: 160
a: [3,37,0,4,1,23]
s: 37
d: [2,134,3,10,3,1,40,99,75]
d: [2,138,3,8,2,2,21,134,57]
s: 138
d: [2,96,2,23,6,0,40,68,27]
d: [2,97,2,21,5,0,18,96]
s: 97
d: [2,95,2,23,6,0,40,68,28]
d: [2,98,2,21,5,0,18,95]
s: 98
d: [2,146,3,10,3,1,40,100,51]
d: [2,151,3,8,2,2,21,146,57]
s: 151
r: 70
r: 64
r: 41
a: [3,40,0,4,1,23]
s: 40
i: [1,22,0,8,2,1,0,0] t42_rewrite1
a: [3,46,0,4,1,22]
s: 46
d: [2,225,5,0,0,4,52,57,32,160,46,100]
e: 225
d: [2,226,3,0,0,4,52,57,32,37,46,107]
e: 226
r: 40
r: 46
s: 41
a: [3,45,0,4,1,22]
s: 45
d: [2,217,4,4,1,3,40,151,41]
a: [3,221,4,4,1,217]
s: 221
d: [2,241,5,0,0,3,52,57,32,221,83]
e: 241
d: [2,243,5,6,2,1,40,221,83]
d: [2,250,5,4,1,2,21,243,57]
d: [2,251,5,0,0,3,21,250,32]
e: 251
i: [1,20,0,8,2,1,0,0] t42_rewrite1
d: [2,234,4,4,1,3,40,41,151]
f: 251,234,226,19,20,21,22,23,24
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rewrite1__t42_rewrite1
% SZS output start Proof for rewrite1__t42_rewrite1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : (((r1_rewrite1(X0,X3,X2) & r5_rewrite1(X0,X1,X2)) | (r5_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2))) => r5_rewrite1(X0,X1,X3))) [input t42_rewrite1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : (((r1_rewrite1(X0,X3,X2) & r5_rewrite1(X0,X1,X2)) | (r5_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2))) => r5_rewrite1(X0,X1,X3))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r5_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3)))) [input d7_rewrite1]
4. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : ((r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2)) => r1_rewrite1(X0,X1,X3))) [input t16_rewrite1]
5. ~! [X0] : (v1_relat_1(X0) => ! [X1,X2,X3] : (((r1_rewrite1(X0,X3,X2) & r5_rewrite1(X0,X1,X2)) | (r5_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2))) => r5_rewrite1(X0,X1,X3))) [flattening 2]
6. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r5_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3)))) [flattening 3]
7. ! [X0] : (v1_relat_1(X0) => ! [X1,X2,X3] : ((r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2)) => r1_rewrite1(X0,X1,X3))) [flattening 4]
8. ? [X0] : (? [X1,X2,X3] : (~r5_rewrite1(X0,X1,X3) & ((r1_rewrite1(X0,X3,X2) & r5_rewrite1(X0,X1,X2)) | (r5_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2)))) & v1_relat_1(X0)) [ennf transformation 5]
9. ! [X0] : (! [X1,X2] : (r5_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3))) | ~v1_relat_1(X0)) [ennf transformation 6]
10. ! [X0] : (! [X1,X2,X3] : (r1_rewrite1(X0,X1,X3) | (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [ennf transformation 7]
11. ! [X0] : (! [X1,X2,X3] : (r1_rewrite1(X0,X1,X3) | ~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X2)) | ~v1_relat_1(X0)) [flattening 10]
12. ? [X0] : (? [X1,X2,X3] : (~r5_rewrite1(X0,X1,X3) & ((r1_rewrite1(X0,X3,X2) & r5_rewrite1(X0,X1,X2)) | (r5_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2)))) & v1_relat_1(X0)) => (? [X3,X2,X1] : (~r5_rewrite1(sK0,X1,X3) & ((r1_rewrite1(sK0,X3,X2) & r5_rewrite1(sK0,X1,X2)) | (r5_rewrite1(sK0,X2,X3) & r1_rewrite1(sK0,X1,X2)))) & v1_relat_1(sK0)) [choice axiom]
13. ? [X3,X2,X1] : (~r5_rewrite1(sK0,X1,X3) & ((r1_rewrite1(sK0,X3,X2) & r5_rewrite1(sK0,X1,X2)) | (r5_rewrite1(sK0,X2,X3) & r1_rewrite1(sK0,X1,X2)))) => (~r5_rewrite1(sK0,sK1,sK3) & ((r1_rewrite1(sK0,sK3,sK2) & r5_rewrite1(sK0,sK1,sK2)) | (r5_rewrite1(sK0,sK2,sK3) & r1_rewrite1(sK0,sK1,sK2)))) [choice axiom]
14. (~r5_rewrite1(sK0,sK1,sK3) & ((r1_rewrite1(sK0,sK3,sK2) & r5_rewrite1(sK0,sK1,sK2)) | (r5_rewrite1(sK0,sK2,sK3) & r1_rewrite1(sK0,sK1,sK2)))) & v1_relat_1(sK0) [skolemisation 8,13,12]
15. ! [X0] : (! [X1,X2] : ((r5_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X3))) & (? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3)) | ~r5_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [nnf transformation 9]
16. ! [X0] : (! [X1,X2] : ((r5_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X3))) & (? [X4] : (r1_rewrite1(X0,X2,X4) & r1_rewrite1(X0,X1,X4)) | ~r5_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [rectify 15]
17. ! [X2,X1,X0] : (? [X4] : (r1_rewrite1(X0,X2,X4) & r1_rewrite1(X0,X1,X4)) => (r1_rewrite1(X0,X2,sK4(X0,X1,X2)) & r1_rewrite1(X0,X1,sK4(X0,X1,X2)))) [choice axiom]
18. ! [X0] : (! [X1,X2] : ((r5_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X3))) & ((r1_rewrite1(X0,X2,sK4(X0,X1,X2)) & r1_rewrite1(X0,X1,sK4(X0,X1,X2))) | ~r5_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [skolemisation 16,17]
19. v1_relat_1(sK0) [cnf transformation 14] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
20. r5_rewrite1(sK0,sK1,sK2) | r1_rewrite1(sK0,sK1,sK2) [cnf transformation 14] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
21. r5_rewrite1(sK0,sK1,sK2) | r5_rewrite1(sK0,sK2,sK3) [cnf transformation 14] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
22. r1_rewrite1(sK0,sK3,sK2) | r1_rewrite1(sK0,sK1,sK2) [cnf transformation 14] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
23. r1_rewrite1(sK0,sK3,sK2) | r5_rewrite1(sK0,sK2,sK3) [cnf transformation 14] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
24. ~r5_rewrite1(sK0,sK1,sK3) [cnf transformation 14] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
25. r1_rewrite1(X0,X1,sK4(X0,X1,X2)) | ~r5_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
26. r1_rewrite1(X0,X2,sK4(X0,X1,X2)) | ~r5_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
27. ~r1_rewrite1(X0,X2,X3) | r5_rewrite1(X0,X1,X2) | ~r1_rewrite1(X0,X1,X3) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
28. ~r1_rewrite1(X0,X2,X3) | r1_rewrite1(X0,X1,X3) | ~r1_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 11] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
30. 1 <=> r5_rewrite1(sK0,sK1,sK3) [avatar definition]
32. ~r5_rewrite1(sK0,sK1,sK3) <- (~1) [avatar component clause 30] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
33. ~1 [avatar split clause 24,30]
35. 2 <=> r5_rewrite1(sK0,sK2,sK3) [avatar definition]
37. r5_rewrite1(sK0,sK2,sK3) <- (2) [avatar component clause 35] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
39. 3 <=> r1_rewrite1(sK0,sK3,sK2) [avatar definition]
41. r1_rewrite1(sK0,sK3,sK2) <- (3) [avatar component clause 39] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
42. 2 | 3 [avatar split clause 23,39,35]
44. 4 <=> r1_rewrite1(sK0,sK1,sK2) [avatar definition]
46. r1_rewrite1(sK0,sK1,sK2) <- (4) [avatar component clause 44] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
47. 4 | 3 [avatar split clause 22,39,44]
49. 5 <=> r5_rewrite1(sK0,sK1,sK2) [avatar definition]
51. r5_rewrite1(sK0,sK1,sK2) <- (5) [avatar component clause 49] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
52. 2 | 5 [avatar split clause 21,49,35]
53. 4 | 5 [avatar split clause 20,49,44]
55. 6 <=> v1_relat_1(sK0) [avatar definition]
57. v1_relat_1(sK0) <- (6) [avatar component clause 55] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
58. 6 [avatar split clause 19,55]
60. r5_rewrite1(X1,X2,X3) | ~r1_rewrite1(X1,X2,sK4(X1,X3,X4)) | ~v1_relat_1(X1) | ~r5_rewrite1(X1,X3,X4) | ~v1_relat_1(X1) [resolution 27,25] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
61. r5_rewrite1(X5,X6,X7) | ~r1_rewrite1(X5,X6,sK4(X5,X8,X7)) | ~v1_relat_1(X5) | ~r5_rewrite1(X5,X8,X7) | ~v1_relat_1(X5) [resolution 27,26] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
62. ~r1_rewrite1(X5,X6,sK4(X5,X8,X7)) | r5_rewrite1(X5,X6,X7) | ~v1_relat_1(X5) | ~r5_rewrite1(X5,X8,X7) [duplicate literal removal 61] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
63. ~r1_rewrite1(X1,X2,sK4(X1,X3,X4)) | r5_rewrite1(X1,X2,X3) | ~v1_relat_1(X1) | ~r5_rewrite1(X1,X3,X4) [duplicate literal removal 60] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
66. r1_rewrite1(X1,X2,sK4(X1,X3,X4)) | ~r1_rewrite1(X1,X2,X3) | ~v1_relat_1(X1) | ~r5_rewrite1(X1,X3,X4) | ~v1_relat_1(X1) [resolution 28,25] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:3,thDist:-3}
67. r1_rewrite1(X5,X6,sK4(X5,X7,X8)) | ~r1_rewrite1(X5,X6,X8) | ~v1_relat_1(X5) | ~r5_rewrite1(X5,X7,X8) | ~v1_relat_1(X5) [resolution 28,26] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:3,thDist:-3}
68. r1_rewrite1(X5,X6,sK4(X5,X7,X8)) | ~r1_rewrite1(X5,X6,X8) | ~v1_relat_1(X5) | ~r5_rewrite1(X5,X7,X8) [duplicate literal removal 67] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
69. r1_rewrite1(X1,X2,sK4(X1,X3,X4)) | ~r1_rewrite1(X1,X2,X3) | ~v1_relat_1(X1) | ~r5_rewrite1(X1,X3,X4) [duplicate literal removal 66] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
82. r5_rewrite1(X3,X4,X5) | ~v1_relat_1(X3) | ~r5_rewrite1(X3,X5,X4) | ~r5_rewrite1(X3,X5,X4) | ~v1_relat_1(X3) [resolution 63,26] {a:2,w:16,wCS:6400,sine:1,thAx:0,allAx:6,thDist:-6}
83. ~r5_rewrite1(X3,X5,X4) | ~v1_relat_1(X3) | r5_rewrite1(X3,X4,X5) [duplicate literal removal 82] {a:2,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
93. ~r1_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) | ~r5_rewrite1(X0,X3,X2) | r5_rewrite1(X0,X1,X3) | ~v1_relat_1(X0) | ~r5_rewrite1(X0,X3,X2) [resolution 68,63] {a:2,w:20,wCS:8000,sine:1,thAx:0,allAx:7,thDist:-7}
100. ~r5_rewrite1(X0,X3,X2) | ~v1_relat_1(X0) | ~r1_rewrite1(X0,X1,X2) | r5_rewrite1(X0,X1,X3) [duplicate literal removal 93] {a:2,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:7,thDist:-7}
102. ~r1_rewrite1(X4,X5,X6) | ~v1_relat_1(X4) | ~r5_rewrite1(X4,X6,X7) | r5_rewrite1(X4,X5,X7) | ~v1_relat_1(X4) | ~r5_rewrite1(X4,X6,X7) [resolution 69,62] {a:2,w:20,wCS:8000,sine:1,thAx:0,allAx:7,thDist:-7}
107. ~r5_rewrite1(X4,X6,X7) | ~v1_relat_1(X4) | ~r1_rewrite1(X4,X5,X6) | r5_rewrite1(X4,X5,X7) [duplicate literal removal 102] {a:2,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:7,thDist:-7}
146. ~v1_relat_1(sK0) | ~r1_rewrite1(sK0,X0,sK2) | r5_rewrite1(sK0,X0,sK1) <- (5) [resolution 100,51] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
151. ~r1_rewrite1(sK0,X0,sK2) | r5_rewrite1(sK0,X0,sK1) <- (5, 6) [subsumption resolution 146,57] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
219. 12 <=> r5_rewrite1(sK0,sK3,sK1) [avatar definition]
221. r5_rewrite1(sK0,sK3,sK1) <- (12) [avatar component clause 219] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
226. $false <- (~1, 2, 4, 6) [unit resulting resolution 57,32,37,46,107] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
232. 1 | ~2 | ~4 | ~6 [avatar contradiction clause 226]
234. r5_rewrite1(sK0,sK3,sK1) <- (3, 5, 6) [resolution 41,151] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
238. 12 | ~3 | ~5 | ~6 [avatar split clause 234,55,49,39,219]
243. ~v1_relat_1(sK0) | r5_rewrite1(sK0,sK1,sK3) <- (12) [resolution 221,83] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
250. r5_rewrite1(sK0,sK1,sK3) <- (6, 12) [subsumption resolution 243,57] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
251. $false <- (~1, 6, 12) [subsumption resolution 250,32] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
252. 1 | ~6 | ~12 [avatar contradiction clause 251]
256. $false [avatar sat refutation 33,42,47,52,53,58,232,238,252] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:258,thDist:-258}
% SZS output end Proof for rewrite1__t42_rewrite1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 3
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 10
% Generated clauses: 187
% Active clauses: 37
% Passive clauses: 69
% Final active clauses: 32
% Final passive clauses: 32
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Duplicate literals: 40
% Fw subsumption resolutions: 38
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 31
% 
% >>> Generating Inferences
% Binary resolution: 94
% Unit resulting resolution: 3
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 2
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

