i: [1,93,0,2,1,1,0,0] t1_net_1
a: [3,156,0,2,1,93]
s: 156
i: [1,97,0,4,1,1,0,0] t1_net_1
a: [3,136,0,4,1,97]
s: 136
i: [1,96,0,4,1,1,0,0] t1_net_1
a: [3,141,0,4,1,96]
s: 141
i: [1,95,0,4,1,1,0,0] t1_net_1
a: [3,146,0,4,1,95]
s: 146
i: [1,94,0,4,1,1,0,0] t1_net_1
a: [3,151,0,4,1,94]
s: 151
i: [1,98,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,161,0,2,1,98]
s: 161
i: [1,99,0,2,1,0,0,1] fc1_xboole_0
a: [3,166,0,2,1,99]
s: 166
i: [1,124,0,2,1,0,0,2] rc2_xboole_0
a: [3,176,0,2,1,124]
s: 176
i: [1,125,0,2,1,0,0,1] existence_l1_petri
a: [3,181,0,2,1,125]
s: 181
i: [1,126,0,2,1,0,0,2] rc1_xboole_0
a: [3,186,0,2,1,126]
s: 186
i: [1,128,0,2,1,0,0,255] rc1_relat_1
a: [3,191,0,2,1,128]
s: 191
i: [1,127,0,2,1,0,0,255] rc1_relat_1
a: [3,196,0,2,1,127]
s: 196
i: [1,129,0,2,1,0,0,255] rc2_relat_1
a: [3,201,0,2,1,129]
s: 201
i: [1,100,0,3,1,0,0,255] d2_xboole_0
a: [3,171,0,3,1,100]
s: 171
r: 161
i: [1,103,0,4,2,0,0,255] cc1_relat_1
s: 103
i: [1,117,0,4,1,0,0,1] existence_m1_subset_1
s: 117
d: [2,204,1,2,1,1,40,103,186]
a: [3,209,1,2,1,204]
s: 209
i: [1,101,0,5,1,0,0,1] t1_boole
s: 101
d: [2,205,1,2,1,1,40,103,166]
a: [3,214,1,2,1,205]
s: 214
i: [1,104,0,5,2,0,0,1] t6_boole
s: 104
i: [1,110,0,5,2,0,0,2] t7_boole
s: 110
i: [1,118,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 118
i: [1,107,0,6,2,0,0,1] t1_subset
s: 107
i: [1,120,0,6,2,0,0,2] fc4_xboole_0
s: 120
i: [1,121,0,6,2,0,0,2] fc5_xboole_0
s: 121
i: [1,122,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 122
i: [1,109,0,7,3,0,0,1] t8_boole
s: 109
i: [1,119,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 119
d: [2,224,1,5,1,0,44,119,101]
s: 224
i: [1,108,0,8,3,0,0,1] t2_subset
s: 108
i: [1,123,0,8,3,0,0,255] fc3_relat_1
s: 123
d: [2,239,1,7,2,1,40,108,146]
a: [3,248,1,3,1,239]
s: 248
d: [2,256,2,0,0,3,52,166,151,248,109]
e: 256
d: [2,257,2,0,0,3,52,166,151,248,109]
e: 257
d: [2,258,2,0,0,2,52,151,248,104]
e: 258
r: 248
i: [1,130,0,8,2,0,0,2] d3_xboole_0
s: 130
a: [3,247,1,3,1,239]
s: 247
i: [1,131,0,8,2,0,0,2] d3_xboole_0
s: 131
i: [1,102,0,10,2,0,0,1] d3_net_1
s: 102
a: [3,244,1,4,1,239]
s: 244
i: [1,132,0,11,3,0,0,2] d3_xboole_0
s: 132
d: [2,284,2,4,1,1,40,244,122]
a: [3,290,2,4,1,284]
s: 290
i: [1,115,0,17,3,0,0,2] d3_xboole_0
s: 115
i: [1,116,0,17,3,0,0,2] d3_xboole_0
s: 116
d: [2,225,1,5,1,0,44,119,101]
s: 225
i: [1,114,0,23,4,0,0,2] d3_xboole_0
s: 114
d: [2,226,1,5,1,0,44,101,119]
s: 226
d: [2,229,1,5,1,0,44,101,119]
s: 229
d: [2,240,1,6,2,0,40,108,117]
s: 240
d: [2,266,1,6,2,0,44,130,101]
s: 266
d: [2,216,1,7,3,0,44,104,104]
s: 216
d: [2,278,2,6,2,0,44,131,224]
s: 278
d: [2,218,1,7,2,0,44,101,104]
s: 218
d: [2,265,1,7,2,0,40,130,110]
s: 265
d: [2,338,2,5,2,0,44,265,101]
a: [3,350,2,2,1,338]
s: 350
r: 196
r: 176
r: 120
r: 121
r: 265
r: 247
r: 103
r: 104
r: 110
r: 109
r: 218
d: [2,354,3,0,0,2,40,350,166]
e: 354
d: [2,353,3,0,0,2,40,350,186]
e: 353
d: [2,352,3,0,0,2,52,186,350]
e: 352
r: 350
s: 176
s: 247
s: 196
s: 103
a: [3,347,2,3,1,338]
s: 347
r: 266
s: 104
s: 110
s: 120
s: 121
d: [2,363,3,5,2,1,44,347,104]
s: 363
s: 109
d: [2,375,1,6,2,0,44,120,119]
s: 375
d: [2,376,1,6,2,0,44,120,119]
s: 376
s: 218
d: [2,317,2,6,2,0,40,240,122]
s: 317
s: 265
d: [2,273,1,7,2,0,40,131,110]
s: 273
d: [2,362,1,7,3,0,44,104,104]
s: 362
d: [2,366,1,7,2,0,44,101,104]
s: 366
d: [2,368,1,7,2,0,40,110,131]
s: 368
d: [2,369,1,7,2,0,40,110,130]
s: 369
d: [2,232,2,7,2,0,44,224,104]
s: 232
d: [2,263,1,8,2,0,40,130,122]
s: 263
d: [2,264,1,8,2,0,40,130,107]
s: 264
d: [2,324,2,7,2,0,44,218,119]
s: 324
d: [2,268,1,8,2,0,44,130,119]
s: 268
d: [2,325,2,7,2,0,44,218,119]
s: 325
d: [2,269,1,8,2,0,44,130,119]
s: 269
d: [2,271,1,8,2,0,40,131,122]
s: 271
d: [2,328,2,7,2,0,44,119,218]
s: 328
d: [2,272,1,8,2,0,40,131,107]
s: 272
d: [2,329,2,7,2,0,44,119,218]
s: 329
d: [2,279,1,8,3,0,44,120,102]
s: 279
d: [2,280,1,8,3,0,44,121,102]
s: 280
d: [2,341,2,7,2,0,44,265,119]
s: 341
d: [2,378,1,8,3,0,44,120,102]
s: 378
d: [2,342,2,7,2,0,44,265,119]
s: 342
d: [2,384,1,8,3,0,44,121,102]
s: 384
d: [2,282,1,10,3,0,44,130,102]
s: 282
d: [2,364,2,7,2,0,44,224,104]
s: 364
d: [2,283,1,10,3,0,44,131,102]
s: 283
d: [2,385,2,7,2,0,44,218,119]
s: 385
d: [2,281,1,11,4,0,44,123,102]
s: 281
d: [2,311,1,17,3,0,42,114]
d: [2,315,1,11,2,0,21,311,116]
s: 315
d: [2,386,2,7,2,0,44,218,119]
s: 386
d: [2,312,1,17,3,0,42,114]
d: [2,316,1,11,2,0,21,312,115]
s: 316
d: [2,389,2,7,2,0,44,119,218]
s: 389
d: [2,299,1,14,4,0,44,132,102]
s: 299
d: [2,309,1,19,4,0,40,114,110]
s: 309
d: [2,495,2,10,3,1,40,299,244]
d: [2,500,2,8,2,2,21,495,156]
a: [3,508,2,4,1,500]
s: 508
d: [2,525,3,0,0,2,52,141,508,107]
e: 525
d: [2,529,3,4,1,1,40,508,107]
d: [2,531,3,0,0,2,21,529,141]
e: 531
r: 508
d: [2,371,1,19,4,0,40,110,114]
s: 371
a: [3,504,2,4,1,500]
s: 504
d: [2,533,3,0,0,2,52,136,504,107]
e: 533
d: [2,537,3,4,1,1,40,504,107]
d: [2,539,3,0,0,2,21,537,136]
e: 539
f: 539,531,500,258,239,93,94,95,96,97
% Refutation found. Thanks to Tanya!
% SZS status Theorem for net_1__t1_net_1
% SZS output start Proof for net_1__t1_net_1
1. ! [X0] : ! [X1] : (l1_petri(X1) => ~(~m1_subset_1(X0,u4_struct_0(X1)) & ~m1_subset_1(X0,u1_struct_0(X1)) & m1_subset_1(X0,k2_net_1(X1)) & k2_net_1(X1) != k1_xboole_0)) [input t1_net_1]
2. ~! [X0] : ! [X1] : (l1_petri(X1) => ~(~m1_subset_1(X0,u4_struct_0(X1)) & ~m1_subset_1(X0,u1_struct_0(X1)) & m1_subset_1(X0,k2_net_1(X1)) & k2_net_1(X1) != k1_xboole_0)) [negated conjecture 1]
18. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
35. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
36. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
39. ! [X0] : (l1_petri(X0) => k2_net_1(X0) = k2_xboole_0(u1_struct_0(X0),u4_struct_0(X0))) [input d3_net_1]
40. ! [X0] : ! [X1] : ! [X2] : (k2_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) | r2_hidden(X3,X0)))) [input d3_xboole_0]
41. ~! [X0,X1] : (l1_petri(X1) => ~(~m1_subset_1(X0,u4_struct_0(X1)) & ~m1_subset_1(X0,u1_struct_0(X1)) & m1_subset_1(X0,k2_net_1(X1)) & k2_net_1(X1) != k1_xboole_0)) [flattening 2]
42. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 35]
43. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 18]
46. ! [X0,X1,X2] : (k2_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) | r2_hidden(X3,X0)))) [flattening 40]
55. ? [X0,X1] : ((~m1_subset_1(X0,u4_struct_0(X1)) & ~m1_subset_1(X0,u1_struct_0(X1)) & m1_subset_1(X0,k2_net_1(X1)) & k2_net_1(X1) != k1_xboole_0) & l1_petri(X1)) [ennf transformation 41]
56. ? [X0,X1] : (~m1_subset_1(X0,u4_struct_0(X1)) & ~m1_subset_1(X0,u1_struct_0(X1)) & m1_subset_1(X0,k2_net_1(X1)) & k2_net_1(X1) != k1_xboole_0 & l1_petri(X1)) [flattening 55]
57. ! [X0] : (k2_net_1(X0) = k2_xboole_0(u1_struct_0(X0),u4_struct_0(X0)) | ~l1_petri(X0)) [ennf transformation 39]
59. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 36]
64. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 42]
65. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 43]
66. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 65]
74. ? [X0,X1] : (~m1_subset_1(X0,u4_struct_0(X1)) & ~m1_subset_1(X0,u1_struct_0(X1)) & m1_subset_1(X0,k2_net_1(X1)) & k2_net_1(X1) != k1_xboole_0 & l1_petri(X1)) => (~m1_subset_1(sK0,u4_struct_0(sK1)) & ~m1_subset_1(sK0,u1_struct_0(sK1)) & m1_subset_1(sK0,k2_net_1(sK1)) & k1_xboole_0 != k2_net_1(sK1) & l1_petri(sK1)) [choice axiom]
75. ~m1_subset_1(sK0,u4_struct_0(sK1)) & ~m1_subset_1(sK0,u1_struct_0(sK1)) & m1_subset_1(sK0,k2_net_1(sK1)) & k1_xboole_0 != k2_net_1(sK1) & l1_petri(sK1) [skolemisation 56,74]
76. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) | r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) & ~r2_hidden(X3,X0))) & ((r2_hidden(X3,X1) | r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k2_xboole_0(X0,X1) != X2)) [nnf transformation 46]
77. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) & ~r2_hidden(X3,X0))) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | ~r2_hidden(X3,X2))) | k2_xboole_0(X0,X1) != X2)) [flattening 76]
78. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (~r2_hidden(X4,X1) & ~r2_hidden(X4,X0))) & (r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2))) | k2_xboole_0(X0,X1) != X2)) [rectify 77]
79. ! [X2,X1,X0] : (? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2))) => (((~r2_hidden(sK2(X0,X1,X2),X1) & ~r2_hidden(sK2(X0,X1,X2),X0)) | ~r2_hidden(sK2(X0,X1,X2),X2)) & (r2_hidden(sK2(X0,X1,X2),X1) | r2_hidden(sK2(X0,X1,X2),X0) | r2_hidden(sK2(X0,X1,X2),X2)))) [choice axiom]
80. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | (((~r2_hidden(sK2(X0,X1,X2),X1) & ~r2_hidden(sK2(X0,X1,X2),X0)) | ~r2_hidden(sK2(X0,X1,X2),X2)) & (r2_hidden(sK2(X0,X1,X2),X1) | r2_hidden(sK2(X0,X1,X2),X0) | r2_hidden(sK2(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (~r2_hidden(X4,X1) & ~r2_hidden(X4,X0))) & (r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2))) | k2_xboole_0(X0,X1) != X2)) [skolemisation 78,79]
93. l1_petri(sK1) [cnf transformation 75] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
94. k1_xboole_0 != k2_net_1(sK1) [cnf transformation 75] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
95. m1_subset_1(sK0,k2_net_1(sK1)) [cnf transformation 75] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
96. ~m1_subset_1(sK0,u1_struct_0(sK1)) [cnf transformation 75] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
97. ~m1_subset_1(sK0,u4_struct_0(sK1)) [cnf transformation 75] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
102. k2_net_1(X0) = k2_xboole_0(u1_struct_0(X0),u4_struct_0(X0)) | ~l1_petri(X0) [cnf transformation 57] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 59] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
107. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 64] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
108. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 66] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
111. r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2) | k2_xboole_0(X0,X1) != X2 [cnf transformation 80] {a:0,w:14,wCS:5600,sine:2,thAx:0,allAx:2,thDist:-2}
132. ~r2_hidden(X4,k2_xboole_0(X0,X1)) | r2_hidden(X4,X0) | r2_hidden(X4,X1) [equality resolution 111] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
134. 1 <=> m1_subset_1(sK0,u4_struct_0(sK1)) [avatar definition]
136. ~m1_subset_1(sK0,u4_struct_0(sK1)) <- (~1) [avatar component clause 134] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
137. ~1 [avatar split clause 97,134]
139. 2 <=> m1_subset_1(sK0,u1_struct_0(sK1)) [avatar definition]
141. ~m1_subset_1(sK0,u1_struct_0(sK1)) <- (~2) [avatar component clause 139] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
142. ~2 [avatar split clause 96,139]
144. 3 <=> m1_subset_1(sK0,k2_net_1(sK1)) [avatar definition]
146. m1_subset_1(sK0,k2_net_1(sK1)) <- (3) [avatar component clause 144] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
147. 3 [avatar split clause 95,144]
149. 4 <=> k1_xboole_0 = k2_net_1(sK1) [avatar definition]
151. k1_xboole_0 != k2_net_1(sK1) <- (~4) [avatar component clause 149] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
152. ~4 [avatar split clause 94,149]
154. 5 <=> l1_petri(sK1) [avatar definition]
156. l1_petri(sK1) <- (5) [avatar component clause 154] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
157. 5 [avatar split clause 93,154]
239. v1_xboole_0(k2_net_1(sK1)) | r2_hidden(sK0,k2_net_1(sK1)) <- (3) [resolution 108,146] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
242. 17 <=> r2_hidden(sK0,k2_net_1(sK1)) [avatar definition]
244. r2_hidden(sK0,k2_net_1(sK1)) <- (17) [avatar component clause 242] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
246. 18 <=> v1_xboole_0(k2_net_1(sK1)) [avatar definition]
248. v1_xboole_0(k2_net_1(sK1)) <- (18) [avatar component clause 246] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
249. 17 | 18 | ~3 [avatar split clause 239,144,246,242]
258. $false <- (~4, 18) [unit resulting resolution 151,248,104] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
262. 4 | ~18 [avatar contradiction clause 258]
299. ~r2_hidden(X13,k2_net_1(X12)) | r2_hidden(X13,u1_struct_0(X12)) | r2_hidden(X13,u4_struct_0(X12)) | ~l1_petri(X12) [superposition 132,102] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
495. r2_hidden(sK0,u1_struct_0(sK1)) | r2_hidden(sK0,u4_struct_0(sK1)) | ~l1_petri(sK1) <- (17) [resolution 299,244] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
500. r2_hidden(sK0,u1_struct_0(sK1)) | r2_hidden(sK0,u4_struct_0(sK1)) <- (5, 17) [subsumption resolution 495,156] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
502. 22 <=> r2_hidden(sK0,u4_struct_0(sK1)) [avatar definition]
504. r2_hidden(sK0,u4_struct_0(sK1)) <- (22) [avatar component clause 502] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
506. 23 <=> r2_hidden(sK0,u1_struct_0(sK1)) [avatar definition]
508. r2_hidden(sK0,u1_struct_0(sK1)) <- (23) [avatar component clause 506] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
509. 22 | 23 | ~5 | ~17 [avatar split clause 500,242,154,506,502]
529. m1_subset_1(sK0,u1_struct_0(sK1)) <- (23) [resolution 508,107] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
531. $false <- (~2, 23) [subsumption resolution 529,141] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
532. 2 | ~23 [avatar contradiction clause 531]
537. m1_subset_1(sK0,u4_struct_0(sK1)) <- (22) [resolution 504,107] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
539. $false <- (~1, 22) [subsumption resolution 537,136] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
540. 1 | ~22 [avatar contradiction clause 539]
541. $false [avatar sat refutation 137,142,147,152,157,249,262,509,532,540] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:145,thDist:-145}
% SZS output end Proof for net_1__t1_net_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 39
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 37
% Generated clauses: 361
% Active clauses: 77
% Passive clauses: 166
% Final active clauses: 61
% Final passive clauses: 62
% Main loop iterations started: 104
% 
% >>> Simplifying Inferences
% Duplicate literals: 10
% Fw subsumption resolutions: 5
% Fw demodulations: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 44
% Forward subsumptions: 126
% Backward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 94
% Unit resulting resolution: 6
% Factoring: 3
% Forward superposition: 109
% Backward superposition: 56
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 32
% SAT solver unit clauses: 14
% SAT solver binary clauses: 12
% 
% Memory used [KB]: 5117
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.004 s ( own 0.003 s ) 

