i: [1,114,0,2,1,1,0,0] t51_ordinal3
a: [3,184,0,2,1,114]
s: 184
i: [1,113,0,2,1,1,0,0] t51_ordinal3
a: [3,189,0,2,1,113]
s: 189
i: [1,115,0,3,1,1,0,0] t51_ordinal3
a: [3,179,0,3,1,115]
s: 179
i: [1,116,0,7,1,1,0,0] t51_ordinal3
a: [3,174,0,7,1,116]
s: 174
i: [1,117,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,194,0,2,1,117]
s: 194
i: [1,118,0,2,1,0,0,2] fc1_xboole_0
a: [3,199,0,2,1,118]
s: 199
i: [1,167,0,14,5,0,0,1] d5_ordinal3
a: [3,209,0,2,1,167]
s: 209
i: [1,151,0,7,3,0,0,2] reflexivity_r1_ordinal1
a: [3,216,0,2,1,151]
s: 216
r: 209
d: [2,273,1,0,0,2,40,216,189]
e: 273
d: [2,272,1,0,0,2,40,216,184]
e: 272
d: [2,271,1,0,0,2,52,184,216]
e: 271
r: 216
s: 209
i: [1,157,0,2,1,0,0,2] rc2_xboole_0
a: [3,224,0,2,1,157]
s: 224
i: [1,158,0,2,1,0,0,2] rc4_ordinal1
a: [3,229,0,2,1,158]
s: 229
i: [1,159,0,2,1,0,0,2] rc1_xboole_0
a: [3,234,0,2,1,159]
s: 234
i: [1,160,0,2,1,0,0,1] rc1_ordinal1
a: [3,239,0,2,1,160]
s: 239
i: [1,164,0,2,1,0,0,255] rc2_ordinal1
a: [3,244,0,2,1,164]
s: 244
i: [1,163,0,2,1,0,0,255] rc2_ordinal1
a: [3,249,0,2,1,163]
s: 249
i: [1,162,0,2,1,0,0,255] rc2_ordinal1
a: [3,254,0,2,1,162]
s: 254
i: [1,161,0,2,1,0,0,255] rc2_ordinal1
a: [3,259,0,2,1,161]
s: 259
i: [1,166,0,2,1,0,0,2] rc5_ordinal1
a: [3,264,0,2,1,166]
s: 264
i: [1,165,0,2,1,0,0,2] rc5_ordinal1
a: [3,269,0,2,1,165]
s: 269
i: [1,147,0,3,1,0,0,3] reflexivity_r1_tarski
s: 147
i: [1,119,0,3,1,0,0,255] d2_xboole_0
a: [3,204,0,3,1,119]
s: 204
r: 194
i: [1,121,0,4,2,0,0,2] cc6_ordinal1
s: 121
i: [1,122,0,4,2,0,0,255] cc9_ordinal1
s: 122
d: [2,278,1,2,1,1,40,121,229]
a: [3,283,1,2,1,278]
s: 283
i: [1,123,0,4,2,0,0,2] cc7_ordinal1
s: 123
i: [1,124,0,4,2,0,0,2] cc3_ordinal1
s: 124
d: [2,314,1,0,0,2,52,209,199,124]
e: 314
d: [2,316,1,2,1,1,40,124,199]
d: [2,318,1,0,0,2,21,316,209]
e: 318
r: 209
a: [3,208,0,2,1,167]
s: 208
i: [1,126,0,4,2,0,0,255] cc1_ordinal1
s: 126
d: [2,327,1,2,1,1,40,126,189]
a: [3,340,1,2,1,327]
s: 340
i: [1,127,0,4,2,0,0,255] cc1_ordinal1
s: 127
i: [1,146,0,4,1,0,0,2] existence_m1_subset_1
s: 146
d: [2,328,1,2,1,1,40,126,184]
a: [3,345,1,2,1,328]
s: 345
i: [1,125,0,5,2,0,0,2] t6_boole
s: 125
d: [2,358,1,2,1,1,40,127,189]
a: [3,371,1,2,1,358]
s: 371
i: [1,143,0,5,2,0,0,1] t7_boole
s: 143
a: [3,219,0,5,2,151]
s: 219
d: [2,359,1,2,1,1,40,127,184]
a: [3,376,1,2,1,359]
s: 376
i: [1,134,0,6,3,0,0,255] cc2_ordinal1
s: 134
d: [2,391,1,2,1,1,40,143,179]
a: [3,395,1,2,1,391]
s: 395
i: [1,136,0,6,2,0,0,3] d2_ordinal1
s: 136
i: [1,137,0,6,2,0,0,3] d2_ordinal1
s: 137
d: [2,279,1,2,1,1,40,121,264]
a: [3,288,1,2,1,279]
s: 288
i: [1,138,0,6,2,0,0,1] t1_subset
s: 138
d: [2,413,1,3,1,1,40,138,179]
a: [3,418,1,3,1,413]
s: 418
i: [1,148,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 148
i: [1,132,0,7,3,0,0,2] cc2_ordinal2
s: 132
d: [2,420,1,3,1,1,40,148,179]
a: [3,425,1,3,1,420]
s: 425
i: [1,133,0,7,3,0,0,2] cc5_ordinal1
s: 133
d: [2,290,1,2,1,1,40,122,234]
a: [3,295,1,2,1,290]
s: 295
i: [1,140,0,7,2,0,0,3] t3_subset
s: 140
i: [1,141,0,7,2,0,0,3] t3_subset
s: 141
d: [2,291,1,2,1,1,40,122,199]
a: [3,300,1,2,1,291]
s: 300
i: [1,142,0,7,3,0,0,1] t8_boole
s: 142
d: [2,302,1,2,1,1,40,123,234]
a: [3,307,1,2,1,302]
s: 307
i: [1,120,0,8,3,0,0,255] cc10_ordinal1
s: 120
i: [1,135,0,8,3,0,0,3] d2_ordinal1
s: 135
d: [2,437,1,5,2,1,40,135,179]
d: [2,439,1,3,1,2,21,437,345]
a: [3,443,1,3,1,439]
s: 443
i: [1,139,0,8,3,0,0,1] t2_subset
s: 139
d: [2,303,1,2,1,1,40,123,199]
a: [3,312,1,2,1,303]
s: 312
i: [1,149,0,8,3,0,0,1] fc5_ordinal2
s: 149
i: [1,150,0,8,3,0,0,1] fc5_ordinal2
s: 150
d: [2,315,1,2,1,1,40,124,234]
a: [3,324,1,2,1,315]
s: 324
i: [1,152,0,8,3,0,0,1] dt_k5_ordinal3
s: 152
d: [2,326,1,2,1,1,40,126,208]
a: [3,335,1,2,1,326]
s: 335
i: [1,153,0,8,3,0,0,1] dt_k10_ordinal2
s: 153
i: [1,145,0,9,3,0,0,4] t5_subset
s: 145
d: [2,330,1,2,1,1,40,126,239]
a: [3,355,1,2,1,330]
s: 355
i: [1,144,0,10,3,0,0,4] t4_subset
s: 144
d: [2,357,1,2,1,1,40,127,208]
a: [3,366,1,2,1,357]
s: 366
i: [1,154,0,10,4,0,0,2] connectedness_r1_ordinal1
s: 154
i: [1,155,0,10,4,0,0,2] redefinition_r1_ordinal1
s: 155
d: [2,361,1,2,1,1,40,127,239]
a: [3,386,1,2,1,361]
s: 386
i: [1,156,0,10,4,0,0,2] redefinition_r1_ordinal1
s: 156
d: [2,484,2,7,3,1,40,156,443]
d: [2,486,2,5,2,2,21,484,184]
d: [2,487,2,3,1,3,21,486,189]
a: [3,491,2,3,1,487]
s: 491
a: [3,212,0,12,4,167]
s: 212
i: [1,168,0,16,5,0,0,1] d5_ordinal3
s: 168
d: [2,329,2,2,1,1,40,126,283]
a: [3,350,2,2,1,329]
s: 350
i: [1,170,0,18,5,0,0,1] d5_ordinal3
s: 170
d: [2,360,2,2,1,1,40,127,283]
a: [3,381,2,2,1,360]
s: 381
d: [2,497,0,14,4,0,21,170,152]
s: 497
d: [2,499,3,0,0,4,52,184,189,491,174,497]
e: 499
d: [2,501,1,10,4,1,44,174,497]
d: [2,506,1,7,3,1,19,501]
d: [2,509,1,4,2,2,21,506,491]
d: [2,510,1,2,1,3,21,509,189]
d: [2,511,1,0,0,4,21,510,184]
e: 511
f: 511,487,439,328,113,114,115,116
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t51_ordinal3
% SZS output start Proof for ordinal3__t51_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_hidden(X0,X1) => k10_ordinal2(X0,k5_ordinal3(X1,X0)) = X1))) [input t51_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_hidden(X0,X1) => k10_ordinal2(X0,k5_ordinal3(X1,X0)) = X1))) [negated conjecture 1]
5. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
15. ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r2_hidden(X1,X0) => r1_tarski(X1,X0))) [input d2_ordinal1]
17. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => ((~r1_ordinal1(X1,X0) => (k5_ordinal3(X0,X1) = X2 <=> k1_xboole_0 = X2)) & (r1_ordinal1(X1,X0) => (k5_ordinal3(X0,X1) = X2 <=> k10_ordinal2(X1,X2) = X0)))))) [input d5_ordinal3]
21. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k5_ordinal3(X0,X1))) [input dt_k5_ordinal3]
33. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
53. ? [X0] : (? [X1] : ((k10_ordinal2(X0,k5_ordinal3(X1,X0)) != X1 & r2_hidden(X0,X1)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
54. ? [X0] : (? [X1] : (k10_ordinal2(X0,k5_ordinal3(X1,X0)) != X1 & r2_hidden(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 53]
61. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 5]
62. ! [X0] : (! [X1] : (! [X2] : ((((k5_ordinal3(X0,X1) = X2 <=> k1_xboole_0 = X2) | r1_ordinal1(X1,X0)) & ((k5_ordinal3(X0,X1) = X2 <=> k10_ordinal2(X1,X2) = X0) | ~r1_ordinal1(X1,X0))) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 17]
67. ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0))) [ennf transformation 15]
81. ! [X0,X1] : (v3_ordinal1(k5_ordinal3(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 21]
82. ! [X0,X1] : (v3_ordinal1(k5_ordinal3(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 81]
87. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 33]
88. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 87]
89. ? [X0] : (? [X1] : (k10_ordinal2(X0,k5_ordinal3(X1,X0)) != X1 & r2_hidden(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (k10_ordinal2(sK0,k5_ordinal3(X1,sK0)) != X1 & r2_hidden(sK0,X1) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
90. ? [X1] : (k10_ordinal2(sK0,k5_ordinal3(X1,sK0)) != X1 & r2_hidden(sK0,X1) & v3_ordinal1(X1)) => (sK1 != k10_ordinal2(sK0,k5_ordinal3(sK1,sK0)) & r2_hidden(sK0,sK1) & v3_ordinal1(sK1)) [choice axiom]
91. (sK1 != k10_ordinal2(sK0,k5_ordinal3(sK1,sK0)) & r2_hidden(sK0,sK1) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 54,90,89]
92. ! [X0] : (! [X1] : (! [X2] : (((((k5_ordinal3(X0,X1) = X2 | k1_xboole_0 != X2) & (k1_xboole_0 = X2 | k5_ordinal3(X0,X1) != X2)) | r1_ordinal1(X1,X0)) & (((k5_ordinal3(X0,X1) = X2 | k10_ordinal2(X1,X2) != X0) & (k10_ordinal2(X1,X2) = X0 | k5_ordinal3(X0,X1) != X2)) | ~r1_ordinal1(X1,X0))) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [nnf transformation 62]
93. ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0)) | ~v1_ordinal1(X0))) [nnf transformation 67]
94. ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0))) [rectify 93]
95. ! [X0] : (? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0)) => (~r1_tarski(sK2(X0),X0) & r2_hidden(sK2(X0),X0))) [choice axiom]
96. ! [X0] : ((v1_ordinal1(X0) | (~r1_tarski(sK2(X0),X0) & r2_hidden(sK2(X0),X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0))) [skolemisation 94,95]
100. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 88]
113. v3_ordinal1(sK0) [cnf transformation 91] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
114. v3_ordinal1(sK1) [cnf transformation 91] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
115. r2_hidden(sK0,sK1) [cnf transformation 91] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
116. sK1 != k10_ordinal2(sK0,k5_ordinal3(sK1,sK0)) [cnf transformation 91] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
126. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 61] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
128. k10_ordinal2(X1,X2) = X0 | k5_ordinal3(X0,X1) != X2 | ~r1_ordinal1(X1,X0) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 92] {a:0,w:19,wCS:7600,sine:1,thAx:0,allAx:1,thDist:-1}
135. ~r2_hidden(X2,X0) | r1_tarski(X2,X0) | ~v1_ordinal1(X0) [cnf transformation 96] {a:0,w:8,wCS:3200,nSel:1,sine:3,thAx:0,allAx:2,thDist:-2}
152. v3_ordinal1(k5_ordinal3(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 82] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
156. ~r1_tarski(X0,X1) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 100] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
170. k10_ordinal2(X1,k5_ordinal3(X0,X1)) = X0 | ~r1_ordinal1(X1,X0) | ~v3_ordinal1(k5_ordinal3(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [equality resolution 128] {a:0,w:18,wCS:7200,sine:1,thAx:0,allAx:1,thDist:-1}
172. 1 <=> sK1 = k10_ordinal2(sK0,k5_ordinal3(sK1,sK0)) [avatar definition]
174. sK1 != k10_ordinal2(sK0,k5_ordinal3(sK1,sK0)) <- (~1) [avatar component clause 172] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
175. ~1 [avatar split clause 116,172]
177. 2 <=> r2_hidden(sK0,sK1) [avatar definition]
179. r2_hidden(sK0,sK1) <- (2) [avatar component clause 177] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
180. 2 [avatar split clause 115,177]
182. 3 <=> v3_ordinal1(sK1) [avatar definition]
184. v3_ordinal1(sK1) <- (3) [avatar component clause 182] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
185. 3 [avatar split clause 114,182]
187. 4 <=> v3_ordinal1(sK0) [avatar definition]
189. v3_ordinal1(sK0) <- (4) [avatar component clause 187] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
190. 4 [avatar split clause 113,187]
328. v1_ordinal1(sK1) <- (3) [resolution 126,184] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
343. 31 <=> v1_ordinal1(sK1) [avatar definition]
345. v1_ordinal1(sK1) <- (31) [avatar component clause 343] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
346. 31 | ~3 [avatar split clause 328,182,343]
437. r1_tarski(sK0,sK1) | ~v1_ordinal1(sK1) <- (2) [resolution 135,179] {a:1,w:5,wCS:500,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
439. r1_tarski(sK0,sK1) <- (2, 31) [subsumption resolution 437,345] {a:1,w:3,wCS:300,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
441. 44 <=> r1_tarski(sK0,sK1) [avatar definition]
443. r1_tarski(sK0,sK1) <- (44) [avatar component clause 441] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
444. 44 | ~2 | ~31 [avatar split clause 439,343,177,441]
484. r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (44) [resolution 156,443] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
486. r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK0) <- (3, 44) [subsumption resolution 484,184] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
487. r1_ordinal1(sK0,sK1) <- (3, 4, 44) [subsumption resolution 486,189] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
489. 47 <=> r1_ordinal1(sK0,sK1) [avatar definition]
491. r1_ordinal1(sK0,sK1) <- (47) [avatar component clause 489] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
492. 47 | ~3 | ~4 | ~44 [avatar split clause 487,441,187,182,489]
497. k10_ordinal2(X1,k5_ordinal3(X0,X1)) = X0 | ~r1_ordinal1(X1,X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [subsumption resolution 170,152] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
501. sK1 != sK1 | ~r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK1) <- (~1) [superposition 174,497] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
506. ~r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK1) <- (~1) [trivial inequality removal 501] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
509. ~v3_ordinal1(sK0) | ~v3_ordinal1(sK1) <- (~1, 47) [subsumption resolution 506,491] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
510. ~v3_ordinal1(sK1) <- (~1, 4, 47) [subsumption resolution 509,189] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
511. $false <- (~1, 3, 4, 47) [subsumption resolution 510,184] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
512. 1 | ~3 | ~4 | ~47 [avatar contradiction clause 511]
514. $false [avatar sat refutation 175,180,185,190,346,444,492,512] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:252,thDist:-252}
% SZS output end Proof for ordinal3__t51_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 42
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 54
% Generated clauses: 209
% Active clauses: 77
% Passive clauses: 109
% Final active clauses: 73
% Final passive clauses: 29
% Main loop iterations started: 80
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Trivial inequalities: 1
% Fw subsumption resolutions: 9
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Forward subsumptions: 24
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 68
% Unit resulting resolution: 3
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 56
% SAT solver unit clauses: 17
% SAT solver binary clauses: 33
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

