i: [1,59,0,2,1,1,0,0] t46_subset_1
a: [3,93,0,2,1,59]
s: 93
i: [1,58,0,2,1,1,0,0] t46_subset_1
a: [3,98,0,2,1,58]
s: 98
i: [1,60,0,7,2,1,0,0] t46_subset_1
s: 60
i: [1,61,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,103,0,2,1,61]
s: 103
i: [1,62,0,2,1,0,0,255] fc1_xboole_0
a: [3,108,0,2,1,62]
s: 108
i: [1,82,0,2,1,0,0,1] rc2_zfmisc_1
a: [3,118,0,2,1,82]
s: 118
i: [1,83,0,2,1,0,0,1] rc2_xboole_0
a: [3,123,0,2,1,83]
s: 123
i: [1,84,0,2,1,0,0,1] rc1_xboole_0
a: [3,128,0,2,1,84]
s: 128
i: [1,86,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,133,0,2,1,86]
s: 133
i: [1,85,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,138,0,2,1,85]
s: 138
i: [1,64,0,3,1,0,0,1] fc2_xboole_0
s: 64
i: [1,65,0,3,1,0,0,1] fc2_zfmisc_1
s: 65
i: [1,63,0,3,1,0,0,255] d2_xboole_0
a: [3,113,0,3,1,63]
s: 113
r: 103
i: [1,66,0,4,2,0,0,1] cc2_zfmisc_1
s: 66
i: [1,67,0,4,2,0,0,1] cc1_zfmisc_1
s: 67
i: [1,80,0,4,1,0,0,1] existence_m1_subset_1
s: 80
d: [2,141,1,2,1,1,40,66,128]
a: [3,146,1,2,1,141]
s: 146
i: [1,88,0,4,1,0,0,2] d1_tarski
s: 88
d: [2,142,1,2,1,1,40,66,108]
a: [3,151,1,2,1,142]
s: 151
i: [1,68,0,5,2,0,0,255] t6_boole
s: 68
i: [1,79,0,5,2,0,0,2] t7_boole
s: 79
i: [1,81,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 81
i: [1,72,0,7,3,0,0,1] d1_subset_1
s: 72
d: [2,157,1,4,1,0,40,81,88]
s: 157
i: [1,73,0,7,3,0,0,1] d1_subset_1
s: 73
i: [1,78,0,7,3,0,0,1] t8_boole
s: 78
d: [2,158,1,5,2,0,40,72,80]
s: 158
i: [1,89,0,7,2,0,0,2] d1_tarski
s: 89
d: [2,161,2,5,2,0,40,158,66]
s: 161
i: [1,70,0,8,3,0,0,1] d1_subset_1
s: 70
i: [1,71,0,8,3,0,0,1] d1_subset_1
s: 71
d: [2,165,0,6,2,0,21,71,79]
s: 165
i: [1,69,0,9,3,0,0,1] t131_zfmisc_1
s: 69
d: [2,172,1,11,4,0,44,60,69]
s: 172
i: [1,76,0,14,3,0,0,2] d1_tarski
s: 76
i: [1,77,0,14,3,0,0,2] d1_tarski
s: 77
d: [2,174,2,8,3,0,47,172]
d: [2,175,2,6,2,1,21,174,93]
d: [2,176,2,4,1,2,21,175,98]
a: [3,180,2,4,1,176]
s: 180
d: [2,166,1,4,1,0,40,165,88]
s: 166
d: [2,163,1,6,2,0,40,70,80]
s: 163
d: [2,153,1,7,3,0,44,68,68]
s: 153
d: [2,169,1,8,3,0,44,88,69]
s: 169
d: [2,193,2,8,2,0,40,163,89]
d: [2,194,2,5,1,0,21,193,64]
s: 194
d: [2,167,1,11,4,0,44,89,69]
d: [2,173,1,9,3,0,21,167,79]
s: 173
d: [2,209,3,8,3,1,44,180,173]
d: [2,214,3,6,2,2,21,209,93]
d: [2,215,3,3,1,2,21,214,165]
s: 215
d: [2,216,4,0,0,3,52,98,80,215,70]
e: 216
d: [2,217,4,0,0,3,52,98,215,163]
e: 217
d: [2,218,4,0,0,3,52,98,93,215,169]
e: 218
d: [2,219,4,4,2,2,40,215,169]
d: [2,225,4,2,1,2,21,219,93]
d: [2,226,4,0,0,3,21,225,98]
e: 226
d: [2,221,4,2,1,2,40,215,163]
d: [2,228,4,0,0,3,21,221,98]
e: 228
f: 228,176,58,59
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t46_subset_1
% SZS output start Proof for subset_1__t46_subset_1
1. ! [X0] : ((v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) => ? [X1] : (k1_tarski(X1) = X0 & m1_subset_1(X1,X0))) [input t46_subset_1]
2. ~! [X0] : ((v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) => ? [X1] : (k1_tarski(X1) = X0 & m1_subset_1(X1,X0))) [negated conjecture 1]
6. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
7. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
13. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
21. ! [X0] : ~(! [X1] : k1_tarski(X1) != X0 & v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) [input t131_zfmisc_1]
23. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
25. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 6]
26. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 7]
28. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 23]
29. ? [X0] : (! [X1] : (k1_tarski(X1) != X0 | ~m1_subset_1(X1,X0)) & (v1_zfmisc_1(X0) & ~v1_xboole_0(X0))) [ennf transformation 2]
30. ? [X0] : (! [X1] : (k1_tarski(X1) != X0 | ~m1_subset_1(X1,X0)) & v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) [flattening 29]
34. ! [X0] : (? [X1] : k1_tarski(X1) = X0 | ~v1_zfmisc_1(X0) | v1_xboole_0(X0)) [ennf transformation 21]
35. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 25]
37. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 28]
39. ? [X0] : (! [X1] : (k1_tarski(X1) != X0 | ~m1_subset_1(X1,X0)) & v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) => (! [X1] : (k1_tarski(X1) != sK0 | ~m1_subset_1(X1,sK0)) & v1_zfmisc_1(sK0) & ~v1_xboole_0(sK0)) [choice axiom]
40. ! [X1] : (k1_tarski(X1) != sK0 | ~m1_subset_1(X1,sK0)) & v1_zfmisc_1(sK0) & ~v1_xboole_0(sK0) [skolemisation 30,39]
41. ! [X0] : (? [X1] : k1_tarski(X1) = X0 => k1_tarski(sK1(X0)) = X0) [choice axiom]
42. ! [X0] : (k1_tarski(sK1(X0)) = X0 | ~v1_zfmisc_1(X0) | v1_xboole_0(X0)) [skolemisation 34,41]
43. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 35]
44. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 26]
45. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 44]
46. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) [choice axiom]
47. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 45,46]
48. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK3(X0),X0)) [choice axiom]
49. ! [X0] : m1_subset_1(sK3(X0),X0) [skolemisation 13,48]
58. ~v1_xboole_0(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. v1_zfmisc_1(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
60. k1_tarski(X1) != sK0 | ~m1_subset_1(X1,sK0) [cnf transformation 40] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. k1_tarski(sK1(X0)) = X0 | ~v1_zfmisc_1(X0) | v1_xboole_0(X0) [cnf transformation 42] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
70. ~m1_subset_1(X1,X0) | r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 43] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. m1_subset_1(X1,X0) | ~r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 43] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
74. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 47] {a:0,w:10,wCS:4000,sine:2,thAx:0,allAx:2,thDist:-2}
79. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 37] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
80. m1_subset_1(sK3(X0),X0) [cnf transformation 49] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
89. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 74] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
91. 1 <=> v1_zfmisc_1(sK0) [avatar definition]
93. v1_zfmisc_1(sK0) <- (1) [avatar component clause 91] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
94. 1 [avatar split clause 59,91]
96. 2 <=> v1_xboole_0(sK0) [avatar definition]
98. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 96] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
99. ~2 [avatar split clause 58,96]
163. r2_hidden(sK3(X0),X0) | v1_xboole_0(X0) [resolution 70,80] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
165. ~r2_hidden(X1,X0) | m1_subset_1(X1,X0) [subsumption resolution 71,79] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
167. ~r2_hidden(X1,X0) | sK1(X0) = X1 | ~v1_zfmisc_1(X0) | v1_xboole_0(X0) [superposition 89,69] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:4,thDist:-4}
172. sK0 != X6 | ~m1_subset_1(sK1(X6),sK0) | ~v1_zfmisc_1(X6) | v1_xboole_0(X6) [superposition 60,69] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
173. sK1(X0) = X1 | ~r2_hidden(X1,X0) | ~v1_zfmisc_1(X0) [subsumption resolution 167,79] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
174. ~m1_subset_1(sK1(sK0),sK0) | ~v1_zfmisc_1(sK0) | v1_xboole_0(sK0) [equality resolution 172] {a:2,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
175. ~m1_subset_1(sK1(sK0),sK0) | v1_xboole_0(sK0) <- (1) [subsumption resolution 174,93] {a:2,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
176. ~m1_subset_1(sK1(sK0),sK0) <- (1, ~2) [subsumption resolution 175,98] {a:2,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
178. 13 <=> m1_subset_1(sK1(sK0),sK0) [avatar definition]
180. ~m1_subset_1(sK1(sK0),sK0) <- (~13) [avatar component clause 178] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
181. ~13 | ~1 | 2 [avatar split clause 176,96,91,178]
209. ~m1_subset_1(X4,sK0) | ~r2_hidden(X4,sK0) | ~v1_zfmisc_1(sK0) <- (~13) [superposition 180,173] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
214. ~m1_subset_1(X4,sK0) | ~r2_hidden(X4,sK0) <- (1, ~13) [subsumption resolution 209,93] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
215. ~r2_hidden(X4,sK0) <- (1, ~13) [subsumption resolution 214,165] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
221. v1_xboole_0(sK0) <- (1, ~13) [resolution 215,163] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
228. $false <- (1, ~2, ~13) [subsumption resolution 221,98] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
229. ~1 | 2 | 13 [avatar contradiction clause 228]
230. $false [avatar sat refutation 94,99,181,229] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:108,thDist:-108}
% SZS output end Proof for subset_1__t46_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 112
% Active clauses: 41
% Passive clauses: 57
% Final active clauses: 40
% Final passive clauses: 13
% Main loop iterations started: 44
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 13
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Equational tautologies: 1
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 29
% Unit resulting resolution: 3
% Forward superposition: 2
% Backward superposition: 15
% Self superposition: 2
% Equality resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 10
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

