i: [1,67,0,3,1,1,0,0] t70_zfmisc_1
a: [3,112,0,3,1,67]
s: 112
i: [1,68,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,122,0,2,1,68]
s: 122
i: [1,69,0,2,1,0,0,255] fc1_xboole_0
a: [3,127,0,2,1,69]
s: 127
i: [1,100,0,2,1,0,0,2] rc2_xboole_0
a: [3,137,0,2,1,100]
s: 137
i: [1,101,0,2,1,0,0,2] rc1_xboole_0
a: [3,142,0,2,1,101]
s: 142
i: [1,66,0,10,1,1,0,0] t70_zfmisc_1
a: [3,117,0,10,1,66]
s: 117
i: [1,92,0,3,1,0,0,2] reflexivity_r1_tarski
s: 92
i: [1,93,0,3,1,0,0,1] reflexivity_r3_xboole_0
s: 93
i: [1,102,0,3,1,0,0,2] d10_xboole_0
s: 102
i: [1,103,0,3,1,0,0,2] d10_xboole_0
s: 103
i: [1,70,0,3,1,0,0,255] d2_xboole_0
a: [3,132,0,3,1,70]
s: 132
r: 122
i: [1,71,0,5,1,0,0,255] t1_boole
s: 71
i: [1,72,0,5,2,0,0,255] t6_boole
s: 72
i: [1,73,0,5,1,0,0,2] t7_xboole_1
s: 73
i: [1,85,0,5,2,0,0,2] t7_boole
s: 85
i: [1,94,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 94
d: [2,148,1,7,1,1,44,73,117]
a: [3,153,1,7,1,148]
s: 153
i: [1,78,0,6,2,0,0,1] d9_xboole_0
s: 78
d: [2,158,2,7,1,1,40,78,153]
a: [3,162,2,7,1,158]
s: 162
i: [1,79,0,6,2,0,0,1] d9_xboole_0
s: 79
i: [1,96,0,6,2,0,0,1] fc5_xboole_0
s: 96
d: [2,172,1,8,2,1,44,96,117]
a: [3,182,1,5,1,172]
s: 182
i: [1,97,0,6,2,0,0,1] fc4_xboole_0
s: 97
d: [2,166,2,7,1,1,40,79,153]
a: [3,170,2,7,1,166]
s: 170
i: [1,98,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 98
i: [1,99,0,6,2,0,0,1] symmetry_r3_xboole_0
s: 99
d: [2,157,1,5,1,0,40,78,73]
s: 157
i: [1,84,0,7,3,0,0,2] t8_boole
s: 84
d: [2,165,1,5,1,0,40,79,73]
s: 165
i: [1,95,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 95
i: [1,104,0,7,2,0,0,1] d1_zfmisc_1
s: 104
d: [2,198,1,5,1,0,44,95,71]
s: 198
i: [1,105,0,7,2,0,0,1] d1_zfmisc_1
s: 105
d: [2,217,2,3,1,0,44,73,198]
s: 217
i: [1,106,0,8,2,0,0,2] d3_xboole_0
s: 106
i: [1,107,0,8,2,0,0,2] d3_xboole_0
s: 107
d: [2,230,1,10,2,1,44,106,117]
s: 230
i: [1,76,0,9,3,0,0,2] d10_xboole_0
s: 76
d: [2,249,2,14,2,1,40,76,153]
a: [3,254,2,7,1,249]
s: 254
i: [1,77,0,9,3,0,0,1] d9_xboole_0
s: 77
i: [1,108,0,11,3,0,0,2] d3_xboole_0
s: 108
d: [2,244,2,9,2,1,40,230,105]
s: 244
i: [1,82,0,14,3,0,0,1] d1_zfmisc_1
s: 82
d: [2,276,3,9,2,1,40,244,79]
s: 276
i: [1,83,0,14,3,0,0,1] d1_zfmisc_1
s: 83
i: [1,90,0,17,3,0,0,2] d3_xboole_0
s: 90
d: [2,277,3,9,2,1,40,244,78]
s: 277
i: [1,91,0,17,3,0,0,2] d3_xboole_0
s: 91
d: [2,282,4,9,2,1,40,276,99]
s: 282
i: [1,89,0,23,4,0,0,2] d3_xboole_0
s: 89
d: [2,238,1,10,2,1,44,107,117]
s: 238
d: [2,308,2,9,2,1,40,238,105]
s: 308
d: [2,268,1,14,3,1,44,108,117]
s: 268
d: [2,314,3,9,2,1,40,308,79]
s: 314
d: [2,199,1,5,1,0,44,95,71]
s: 199
d: [2,200,1,5,1,0,44,71,95]
s: 200
d: [2,315,3,9,2,1,40,308,78]
s: 315
d: [2,201,1,5,1,0,44,73,95]
s: 201
d: [2,328,2,7,1,1,44,201,117]
a: [3,337,2,7,1,328]
s: 337
d: [2,206,1,5,1,0,44,71,95]
s: 206
d: [2,207,1,5,1,0,44,73,95]
s: 207
d: [2,339,3,14,2,1,40,337,76]
a: [3,345,3,7,1,339]
s: 345
d: [2,213,1,6,2,0,40,104,85]
s: 213
d: [2,340,3,7,1,1,40,337,79]
a: [3,354,3,7,1,340]
s: 354
d: [2,231,1,6,2,0,44,106,71]
s: 231
d: [2,145,1,7,3,0,44,72,72]
s: 145
d: [2,341,3,7,1,1,40,337,78]
a: [3,359,3,7,1,341]
s: 359
d: [2,146,1,7,2,0,44,71,72]
s: 146
d: [2,388,2,5,2,1,44,182,146]
a: [3,408,2,2,1,388]
s: 408
a: [3,177,1,3,1,172]
s: 177
d: [2,212,1,7,2,0,40,104,98]
s: 212
d: [2,322,4,9,2,1,40,314,99]
s: 322
d: [2,229,1,7,2,0,40,106,85]
s: 229
d: [2,245,2,10,2,1,40,230,98]
s: 245
d: [2,237,1,7,2,0,40,107,85]
s: 237
d: [2,228,1,8,2,0,40,106,98]
s: 228
d: [2,480,3,14,2,1,40,245,238]
a: [3,493,3,7,1,480]
s: 493
d: [2,233,1,8,2,0,44,106,95]
s: 233
d: [2,516,4,6,1,1,40,493,104]
a: [3,521,4,6,1,516]
s: 521
d: [2,234,1,8,2,0,44,106,95]
s: 234
d: [2,236,1,8,2,0,40,107,98]
s: 236
d: [2,309,2,10,2,1,40,238,98]
s: 309
d: [2,248,1,10,2,0,40,76,73]
s: 248
a: [3,489,3,7,1,480]
s: 489
d: [2,302,1,17,3,0,42,89]
d: [2,306,1,11,2,0,21,302,91]
s: 306
d: [2,303,1,17,3,0,42,89]
d: [2,307,1,11,2,0,21,303,90]
s: 307
d: [2,553,4,6,1,1,40,489,104]
a: [3,558,4,6,1,553]
s: 558
d: [2,279,1,14,3,0,40,82,79]
s: 279
d: [2,529,2,10,2,1,44,236,117]
s: 529
d: [2,280,1,14,3,0,40,82,78]
s: 280
d: [2,278,1,19,4,0,40,82,76]
s: 278
d: [2,461,2,11,2,1,40,212,238]
s: 461
d: [2,299,1,19,4,0,40,89,85]
s: 299
d: [2,594,3,8,1,1,40,461,92]
a: [3,604,3,8,1,594]
s: 604
d: [2,286,1,20,3,0,40,90,104]
s: 286
d: [2,292,1,20,3,0,40,91,104]
s: 292
d: [2,619,4,7,1,1,40,604,104]
a: [3,624,4,7,1,619]
s: 624
d: [2,288,1,23,3,0,40,90,107]
s: 288
d: [2,595,3,10,1,1,40,461,73]
s: 595
d: [2,289,1,23,3,0,40,90,106]
s: 289
d: [2,294,1,23,3,0,40,91,107]
s: 294
d: [2,652,4,9,1,1,40,595,104]
s: 652
d: [2,295,1,23,3,0,40,91,106]
s: 295
d: [2,704,5,9,1,1,44,652,95]
s: 704
d: [2,298,1,23,4,0,40,89,98]
s: 298
d: [2,300,1,27,4,0,40,89,105]
s: 300
d: [2,705,5,9,1,1,44,652,95]
s: 705
d: [2,301,1,39,5,0,40,89,108]
s: 301
d: [2,598,3,10,1,1,40,461,201]
s: 598
d: [2,462,2,11,2,1,40,212,230]
s: 462
d: [2,505,2,12,2,1,40,228,238]
s: 505
d: [2,795,3,10,2,1,44,505,71]
a: [3,804,3,6,1,795]
s: 804
d: [2,506,2,12,2,1,40,228,230]
s: 506
d: [2,781,3,8,1,1,40,462,92]
a: [3,791,3,8,1,781]
s: 791
d: [2,525,2,12,2,1,40,236,238]
s: 525
d: [2,526,2,12,2,1,40,236,230]
s: 526
d: [2,820,4,7,1,1,40,791,104]
a: [3,825,4,7,1,820]
s: 825
d: [2,222,2,3,1,0,44,157,198]
s: 222
d: [2,782,3,10,1,1,40,462,73]
s: 782
d: [2,223,2,3,1,0,44,165,198]
s: 223
d: [2,466,2,5,2,0,44,229,71]
a: [3,474,2,3,1,466]
s: 474
r: 804
r: 231
d: [2,848,4,9,1,1,40,782,104]
s: 848
d: [2,319,2,13,3,1,40,268,104]
s: 319
d: [2,879,3,8,2,1,40,319,73]
a: [3,889,3,4,1,879]
s: 889
a: [3,412,2,3,1,388]
s: 412
d: [2,190,2,5,1,0,40,157,99]
s: 190
d: [2,880,3,8,2,1,40,319,201]
a: [3,898,3,4,1,880]
s: 898
r: 898
d: [2,204,2,5,1,0,44,157,95]
s: 204
a: [3,897,3,4,1,880]
s: 897
d: [2,205,2,5,1,0,44,165,95]
s: 205
d: [2,210,2,5,1,0,44,157,95]
s: 210
d: [2,931,4,3,1,1,40,898,105]
a: [3,936,4,3,1,931]
s: 936
d: [2,211,2,5,1,0,44,165,95]
s: 211
a: [3,902,3,4,1,880]
s: 902
d: [2,326,2,5,1,0,40,201,79]
s: 326
d: [2,327,2,5,1,0,40,201,78]
s: 327
d: [2,882,3,8,2,1,40,319,217]
a: [3,916,3,4,1,882]
s: 916
d: [2,382,2,5,2,0,44,201,146]
s: 382
d: [2,923,4,4,1,1,40,889,98]
a: [3,928,4,4,1,923]
s: 928
d: [2,243,2,6,2,0,44,107,198]
s: 243
d: [2,510,2,6,2,0,44,228,71]
s: 510
d: [2,965,4,4,1,1,40,902,98]
a: [3,970,4,4,1,965]
s: 970
d: [2,535,2,6,2,0,44,236,198]
s: 535
d: [2,881,3,12,2,1,40,319,92]
a: [3,907,3,6,1,881]
s: 907
r: 907
d: [2,552,2,6,2,0,44,248,198]
s: 552
d: [2,287,2,27,3,1,40,90,230]
s: 287
a: [3,906,3,6,1,881]
s: 906
d: [2,293,2,27,3,1,40,91,230]
s: 293
d: [2,1007,4,5,1,1,40,907,105]
a: [3,1024,4,5,1,1007]
s: 1024
d: [2,311,2,27,3,1,40,238,91]
s: 311
d: [2,312,2,27,3,1,40,238,90]
s: 312
d: [2,1080,5,4,1,1,40,1024,382]
a: [3,1085,5,4,1,1080]
s: 1085
d: [2,639,2,27,3,1,44,288,117]
s: 639
a: [3,911,3,6,1,881]
s: 911
r: 906
r: 911
r: 1024
d: [2,667,2,27,3,1,44,289,117]
s: 667
d: [2,688,2,27,3,1,44,294,117]
s: 688
d: [2,932,4,4,1,1,40,898,98]
a: [3,942,4,4,1,932]
s: 942
d: [2,714,2,27,3,1,44,295,117]
s: 714
a: [3,1025,4,5,1,1007]
s: 1025
r: 1025
d: [2,565,2,28,3,1,40,306,268]
s: 565
d: [2,575,2,28,3,1,40,307,268]
s: 575
d: [2,1129,4,4,1,1,40,911,236]
a: [3,1140,4,4,1,1129]
s: 1140
d: [2,214,2,7,2,0,44,198,72]
s: 214
d: [2,1276,3,7,2,1,44,153,214]
a: [3,1323,3,2,1,1276]
s: 1323
d: [2,369,2,7,2,0,44,146,95]
s: 369
d: [2,370,2,7,2,0,44,146,95]
s: 370
s: 1024
d: [2,373,2,7,2,0,44,95,146]
s: 373
d: [2,1131,4,5,1,1,40,911,105]
a: [3,1145,4,5,1,1131]
s: 1145
r: 254
r: 1024
r: 1145
a: [3,258,2,7,1,249]
s: 258
r: 117
r: 153
r: 162
r: 170
r: 182
r: 230
r: 238
r: 245
r: 268
r: 287
r: 293
r: 309
r: 311
r: 312
r: 337
r: 345
r: 354
r: 359
r: 461
r: 462
r: 489
r: 493
r: 505
r: 506
r: 521
r: 525
r: 526
r: 558
r: 565
r: 575
r: 595
r: 598
r: 604
r: 624
r: 652
r: 704
r: 782
r: 791
r: 825
r: 848
r: 258
s: 117
a: [3,807,3,4,1,795]
s: 807
r: 897
r: 308
r: 314
r: 315
d: [2,1607,4,0,0,2,40,807,889]
e: 1607
d: [2,1606,4,0,0,2,40,807,916]
e: 1606
d: [2,1603,4,0,0,2,52,916,807]
e: 1603
d: [2,1604,4,0,0,1,52,92,807,104]
e: 1604
r: 807
s: 182
s: 897
s: 153
s: 230
s: 1024
s: 238
s: 1145
d: [2,1567,1,10,2,1,44,106,117]
s: 1567
d: [2,1568,1,10,2,1,44,107,117]
s: 1568
d: [2,1437,5,10,2,1,40,1145,76]
d: [2,1442,5,5,1,1,21,1437,201]
a: [3,1446,5,5,1,1442]
s: 1446
r: 117
r: 153
r: 182
r: 230
r: 238
r: 244
r: 276
r: 277
r: 319
r: 1024
r: 1085
r: 1145
d: [2,1688,6,3,1,1,44,157,1446]
d: [2,1717,6,0,0,2,21,1688,112]
e: 1717
r: 897
r: 936
r: 1446
d: [2,1450,0,8,1,2,23,117,258]
a: [3,1501,0,8,1,1450]
s: 1501
a: [3,937,4,3,1,931]
s: 937
d: [2,1748,5,0,0,2,52,112,937,79]
e: 1748
d: [2,1750,5,3,1,1,40,937,79]
d: [2,1753,5,0,0,2,21,1750,112]
e: 1753
r: 937
s: 117
s: 182
s: 936
s: 153
s: 897
s: 230
s: 238
d: [2,1474,4,4,1,2,23,521,258]
a: [3,1529,4,4,1,1474]
s: 1529
s: 268
d: [2,1477,4,4,1,2,23,558,258]
a: [3,1534,4,4,1,1477]
s: 1534
d: [2,1569,1,14,3,1,44,108,117]
s: 1569
s: 162
s: 1085
s: 170
d: [2,399,3,7,2,1,44,337,146]
a: [3,443,3,5,1,399]
s: 443
a: [3,253,2,7,1,249]
s: 253
s: 258
r: 153
r: 162
r: 170
r: 182
r: 230
r: 238
r: 253
r: 268
d: [2,1470,3,5,1,2,23,489,258]
a: [3,1519,3,5,1,1470]
s: 1519
s: 337
d: [2,1471,3,5,1,2,23,493,258]
a: [3,1524,3,5,1,1471]
s: 1524
d: [2,389,2,10,3,1,44,230,146]
a: [3,416,2,8,2,389]
s: 416
d: [2,1859,4,0,0,3,52,897,902,416]
e: 1859
d: [2,1861,4,4,1,2,40,416,902]
d: [2,1870,4,0,0,3,21,1861,897]
e: 1870
d: [2,1153,5,10,2,1,40,1025,76]
d: [2,1162,5,5,1,1,21,1153,73]
d: [2,1564,1,7,1,1,44,73,117]
d: [2,1581,2,14,2,1,44,248,117]
d: [2,1661,5,10,2,1,40,1145,76]
d: [2,1666,5,5,1,1,21,1661,201]
d: [2,1843,1,8,2,2,23,230,258]
f: 1870,1843,1753,1717,1666,1581,1564,1162,1131,1007,931,881,880,66,67
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t70_zfmisc_1
% SZS output start Proof for zfmisc_1__t70_zfmisc_1
1. ! [X0] : ! [X1] : (k2_xboole_0(k1_zfmisc_1(X0),k1_zfmisc_1(X1)) = k1_zfmisc_1(k2_xboole_0(X0,X1)) => r3_xboole_0(X0,X1)) [input t70_zfmisc_1]
2. ~! [X0] : ! [X1] : (k2_xboole_0(k1_zfmisc_1(X0),k1_zfmisc_1(X1)) = k1_zfmisc_1(k2_xboole_0(X0,X1)) => r3_xboole_0(X0,X1)) [negated conjecture 1]
4. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
5. ! [X0] : ! [X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [input d10_xboole_0]
6. ! [X0] : ! [X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [input d1_zfmisc_1]
8. ! [X0] : ! [X1] : ! [X2] : (k2_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) | r2_hidden(X3,X0)))) [input d3_xboole_0]
9. ! [X0] : ! [X1] : (r3_xboole_0(X0,X1) <=> (r1_tarski(X1,X0) | r1_tarski(X0,X1))) [input d9_xboole_0]
20. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
26. ! [X0] : ! [X1] : r1_tarski(X0,k2_xboole_0(X0,X1)) [input t7_xboole_1]
28. ~! [X0,X1] : (k2_xboole_0(k1_zfmisc_1(X0),k1_zfmisc_1(X1)) = k1_zfmisc_1(k2_xboole_0(X0,X1)) => r3_xboole_0(X0,X1)) [flattening 2]
29. ! [X0,X1] : r1_tarski(X0,k2_xboole_0(X0,X1)) [flattening 26]
30. ! [X0,X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [flattening 5]
31. ! [X0,X1] : (r3_xboole_0(X0,X1) <=> (r1_tarski(X1,X0) | r1_tarski(X0,X1))) [flattening 9]
32. ! [X0,X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [flattening 6]
35. ! [X0,X1,X2] : (k2_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) | r2_hidden(X3,X0)))) [flattening 8]
36. ! [X0] : r1_tarski(X0,X0) [rectify 20]
39. ? [X0,X1] : (~r3_xboole_0(X0,X1) & k2_xboole_0(k1_zfmisc_1(X0),k1_zfmisc_1(X1)) = k1_zfmisc_1(k2_xboole_0(X0,X1))) [ennf transformation 28]
47. ? [X0,X1] : (~r3_xboole_0(X0,X1) & k2_xboole_0(k1_zfmisc_1(X0),k1_zfmisc_1(X1)) = k1_zfmisc_1(k2_xboole_0(X0,X1))) => (~r3_xboole_0(sK0,sK1) & k2_xboole_0(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) = k1_zfmisc_1(k2_xboole_0(sK0,sK1))) [choice axiom]
48. ~r3_xboole_0(sK0,sK1) & k2_xboole_0(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) = k1_zfmisc_1(k2_xboole_0(sK0,sK1)) [skolemisation 39,47]
49. ! [X0,X1] : ((X0 = X1 | (~r1_tarski(X1,X0) | ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [nnf transformation 30]
50. ! [X0,X1] : ((X0 = X1 | ~r1_tarski(X1,X0) | ~r1_tarski(X0,X1)) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [flattening 49]
51. ! [X0,X1] : ((r3_xboole_0(X0,X1) | (~r1_tarski(X1,X0) & ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) | r1_tarski(X0,X1)) | ~r3_xboole_0(X0,X1))) [nnf transformation 31]
52. ! [X0,X1] : ((r3_xboole_0(X0,X1) | (~r1_tarski(X1,X0) & ~r1_tarski(X0,X1))) & (r1_tarski(X1,X0) | r1_tarski(X0,X1) | ~r3_xboole_0(X0,X1))) [flattening 51]
53. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ~r1_tarski(X2,X0)) & (r1_tarski(X2,X0) | ~r2_hidden(X2,X1))) | k1_zfmisc_1(X0) != X1)) [nnf transformation 32]
54. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [rectify 53]
55. ! [X1,X0] : (? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1))) => ((~r1_tarski(sK2(X0,X1),X0) | ~r2_hidden(sK2(X0,X1),X1)) & (r1_tarski(sK2(X0,X1),X0) | r2_hidden(sK2(X0,X1),X1)))) [choice axiom]
56. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ((~r1_tarski(sK2(X0,X1),X0) | ~r2_hidden(sK2(X0,X1),X1)) & (r1_tarski(sK2(X0,X1),X0) | r2_hidden(sK2(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [skolemisation 54,55]
57. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) | r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) & ~r2_hidden(X3,X0))) & ((r2_hidden(X3,X1) | r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k2_xboole_0(X0,X1) != X2)) [nnf transformation 35]
58. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) & ~r2_hidden(X3,X0))) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | ~r2_hidden(X3,X2))) | k2_xboole_0(X0,X1) != X2)) [flattening 57]
59. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (~r2_hidden(X4,X1) & ~r2_hidden(X4,X0))) & (r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2))) | k2_xboole_0(X0,X1) != X2)) [rectify 58]
60. ! [X2,X1,X0] : (? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2))) => (((~r2_hidden(sK3(X0,X1,X2),X1) & ~r2_hidden(sK3(X0,X1,X2),X0)) | ~r2_hidden(sK3(X0,X1,X2),X2)) & (r2_hidden(sK3(X0,X1,X2),X1) | r2_hidden(sK3(X0,X1,X2),X0) | r2_hidden(sK3(X0,X1,X2),X2)))) [choice axiom]
61. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | (((~r2_hidden(sK3(X0,X1,X2),X1) & ~r2_hidden(sK3(X0,X1,X2),X0)) | ~r2_hidden(sK3(X0,X1,X2),X2)) & (r2_hidden(sK3(X0,X1,X2),X1) | r2_hidden(sK3(X0,X1,X2),X0) | r2_hidden(sK3(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (~r2_hidden(X4,X1) & ~r2_hidden(X4,X0))) & (r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2))) | k2_xboole_0(X0,X1) != X2)) [skolemisation 59,60]
66. k2_xboole_0(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) = k1_zfmisc_1(k2_xboole_0(sK0,sK1)) [cnf transformation 48] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
67. ~r3_xboole_0(sK0,sK1) [cnf transformation 48] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
73. r1_tarski(X0,k2_xboole_0(X0,X1)) [cnf transformation 29] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
76. ~r1_tarski(X1,X0) | X0 = X1 | ~r1_tarski(X0,X1) [cnf transformation 50] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
78. ~r1_tarski(X0,X1) | r3_xboole_0(X0,X1) [cnf transformation 52] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. ~r1_tarski(X1,X0) | r3_xboole_0(X0,X1) [cnf transformation 52] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
80. r1_tarski(X3,X0) | ~r2_hidden(X3,X1) | k1_zfmisc_1(X0) != X1 [cnf transformation 56] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
81. r2_hidden(X3,X1) | ~r1_tarski(X3,X0) | k1_zfmisc_1(X0) != X1 [cnf transformation 56] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
86. r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2) | k2_xboole_0(X0,X1) != X2 [cnf transformation 61] {a:0,w:14,wCS:5600,sine:2,thAx:0,allAx:2,thDist:-2}
88. r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | k2_xboole_0(X0,X1) != X2 [cnf transformation 61] {a:0,w:11,wCS:4400,sine:2,thAx:0,allAx:2,thDist:-2}
92. r1_tarski(X0,X0) [cnf transformation 36] {a:0,w:3,wCS:1200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
95. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. r2_hidden(X3,k1_zfmisc_1(X0)) | ~r1_tarski(X3,X0) [equality resolution 81] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
105. ~r2_hidden(X3,k1_zfmisc_1(X0)) | r1_tarski(X3,X0) [equality resolution 80] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
106. r2_hidden(X4,k2_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) [equality resolution 88] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
108. ~r2_hidden(X4,k2_xboole_0(X0,X1)) | r2_hidden(X4,X0) | r2_hidden(X4,X1) [equality resolution 86] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
110. 1 <=> r3_xboole_0(sK0,sK1) [avatar definition]
112. ~r3_xboole_0(sK0,sK1) <- (~1) [avatar component clause 110] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
113. ~1 [avatar split clause 67,110]
115. 2 <=> k2_xboole_0(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) = k1_zfmisc_1(k2_xboole_0(sK0,sK1)) [avatar definition]
117. k2_xboole_0(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) = k1_zfmisc_1(k2_xboole_0(sK0,sK1)) <- (2) [avatar component clause 115] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
118. 2 [avatar split clause 66,115]
151. 8 <=> r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(k2_xboole_0(sK0,sK1))) [avatar definition]
157. r3_xboole_0(X1,k2_xboole_0(X1,X2)) [resolution 78,73] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
201. r1_tarski(X1,k2_xboole_0(X2,X1)) [superposition 73,95] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
230. r2_hidden(X0,k1_zfmisc_1(k2_xboole_0(sK0,sK1))) | ~r2_hidden(X0,k1_zfmisc_1(sK1)) <- (2) [superposition 106,117] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
248. ~r1_tarski(k2_xboole_0(X1,X2),X1) | k2_xboole_0(X1,X2) = X1 [resolution 76,73] {a:1,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
252. 13 <=> r1_tarski(k1_zfmisc_1(k2_xboole_0(sK0,sK1)),k1_zfmisc_1(sK0)) [avatar definition]
256. 14 <=> k1_zfmisc_1(sK0) = k1_zfmisc_1(k2_xboole_0(sK0,sK1)) [avatar definition]
258. k1_zfmisc_1(sK0) = k1_zfmisc_1(k2_xboole_0(sK0,sK1)) <- (14) [avatar component clause 256] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
268. ~r2_hidden(X0,k1_zfmisc_1(k2_xboole_0(sK0,sK1))) | r2_hidden(X0,k1_zfmisc_1(sK0)) | r2_hidden(X0,k1_zfmisc_1(sK1)) <- (2) [superposition 108,117] {a:1,w:14,wCS:1400,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
319. ~r1_tarski(X2,k2_xboole_0(sK0,sK1)) | r2_hidden(X2,k1_zfmisc_1(sK1)) | r2_hidden(X2,k1_zfmisc_1(sK0)) <- (2) [resolution 268,104] {a:2,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
415. 22 <=> ! [X28] : (r2_hidden(X28,k1_zfmisc_1(sK0)) | ~r2_hidden(X28,k1_zfmisc_1(sK1))) [avatar definition]
416. ~r2_hidden(X28,k1_zfmisc_1(sK1)) | r2_hidden(X28,k1_zfmisc_1(sK0)) <- (22) [avatar component clause 415] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
880. r2_hidden(sK1,k1_zfmisc_1(sK1)) | r2_hidden(sK1,k1_zfmisc_1(sK0)) <- (2) [resolution 319,201] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
881. r2_hidden(k2_xboole_0(sK0,sK1),k1_zfmisc_1(sK1)) | r2_hidden(k2_xboole_0(sK0,sK1),k1_zfmisc_1(sK0)) <- (2) [resolution 319,92] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
896. 46 <=> r2_hidden(sK1,k1_zfmisc_1(sK0)) [avatar definition]
897. ~r2_hidden(sK1,k1_zfmisc_1(sK0)) <- (~46) [avatar component clause 896] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
898. r2_hidden(sK1,k1_zfmisc_1(sK0)) <- (46) [avatar component clause 896] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
900. 47 <=> r2_hidden(sK1,k1_zfmisc_1(sK1)) [avatar definition]
902. r2_hidden(sK1,k1_zfmisc_1(sK1)) <- (47) [avatar component clause 900] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
903. 46 | 47 | ~2 [avatar split clause 880,115,900,896]
905. 48 <=> r2_hidden(k2_xboole_0(sK0,sK1),k1_zfmisc_1(sK0)) [avatar definition]
907. r2_hidden(k2_xboole_0(sK0,sK1),k1_zfmisc_1(sK0)) <- (48) [avatar component clause 905] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
909. 49 <=> r2_hidden(k2_xboole_0(sK0,sK1),k1_zfmisc_1(sK1)) [avatar definition]
911. r2_hidden(k2_xboole_0(sK0,sK1),k1_zfmisc_1(sK1)) <- (49) [avatar component clause 909] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
912. 48 | 49 | ~2 [avatar split clause 881,115,909,905]
931. r1_tarski(sK1,sK0) <- (46) [resolution 898,105] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
935. 53 <=> r1_tarski(sK1,sK0) [avatar definition]
937. r1_tarski(sK1,sK0) <- (53) [avatar component clause 935] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
938. 53 | ~46 [avatar split clause 931,896,935]
1007. r1_tarski(k2_xboole_0(sK0,sK1),sK0) <- (48) [resolution 907,105] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1023. 59 <=> r1_tarski(k2_xboole_0(sK0,sK1),sK0) [avatar definition]
1025. r1_tarski(k2_xboole_0(sK0,sK1),sK0) <- (59) [avatar component clause 1023] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1026. 59 | ~48 [avatar split clause 1007,905,1023]
1131. r1_tarski(k2_xboole_0(sK0,sK1),sK1) <- (49) [resolution 911,105] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1143. 71 <=> r1_tarski(k2_xboole_0(sK0,sK1),sK1) [avatar definition]
1145. r1_tarski(k2_xboole_0(sK0,sK1),sK1) <- (71) [avatar component clause 1143] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1146. 71 | ~49 [avatar split clause 1131,909,1143]
1153. sK0 = k2_xboole_0(sK0,sK1) | ~r1_tarski(sK0,k2_xboole_0(sK0,sK1)) <- (59) [resolution 1025,76] {a:5,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
1158. 73 <=> sK0 = k2_xboole_0(sK0,sK1) [avatar definition]
1162. sK0 = k2_xboole_0(sK0,sK1) <- (59) [subsumption resolution 1153,73] {a:5,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
1163. 73 | ~59 [avatar split clause 1162,1023,1158]
1444. 102 <=> sK1 = k2_xboole_0(sK0,sK1) [avatar definition]
1446. sK1 = k2_xboole_0(sK0,sK1) <- (102) [avatar component clause 1444] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
1448. sK0 != k2_xboole_0(sK0,sK1) | ~r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(k2_xboole_0(sK0,sK1))) | r1_tarski(k1_zfmisc_1(k2_xboole_0(sK0,sK1)),k1_zfmisc_1(sK0)) [theory tautology sat conflict] {a:0,w:19,wCS:7600,sine:255,thAx:0,allAx:1,thDist:-1}
1564. r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(k2_xboole_0(sK0,sK1))) <- (2) [superposition 73,117] {a:1,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
1581. ~r1_tarski(k1_zfmisc_1(k2_xboole_0(sK0,sK1)),k1_zfmisc_1(sK0)) | k1_zfmisc_1(sK0) = k1_zfmisc_1(k2_xboole_0(sK0,sK1)) <- (2) [superposition 248,117] {a:2,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1587. 8 | ~2 [avatar split clause 1564,115,151]
1602. 14 | ~13 | ~2 [avatar split clause 1581,115,252,256]
1661. sK1 = k2_xboole_0(sK0,sK1) | ~r1_tarski(sK1,k2_xboole_0(sK0,sK1)) <- (71) [resolution 1145,76] {a:5,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
1666. sK1 = k2_xboole_0(sK0,sK1) <- (71) [subsumption resolution 1661,201] {a:5,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1667. 102 | ~71 [avatar split clause 1666,1143,1444]
1688. r3_xboole_0(sK0,sK1) <- (102) [superposition 157,1446] {a:6,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
1717. $false <- (~1, 102) [subsumption resolution 1688,112] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
1718. 1 | ~102 [avatar contradiction clause 1717]
1750. r3_xboole_0(sK0,sK1) <- (53) [resolution 937,79] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
1753. $false <- (~1, 53) [subsumption resolution 1750,112] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
1754. 1 | ~53 [avatar contradiction clause 1753]
1843. r2_hidden(X0,k1_zfmisc_1(sK0)) | ~r2_hidden(X0,k1_zfmisc_1(sK1)) <- (2, 14) [backward demodulation 230,258] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1855. 22 | ~2 | ~14 [avatar split clause 1843,256,115,415]
1861. r2_hidden(sK1,k1_zfmisc_1(sK0)) <- (22, 47) [resolution 416,902] {a:4,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1870. $false <- (22, ~46, 47) [subsumption resolution 1861,897] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1871. ~22 | 46 | ~47 [avatar contradiction clause 1870]
1872. $false [avatar sat refutation 113,118,903,912,938,1026,1146,1163,1448,1587,1602,1667,1718,1754,1855,1871] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1744,thDist:-1744}
% SZS output end Proof for zfmisc_1__t70_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 1354
% Active clauses: 172
% Passive clauses: 548
% Final active clauses: 90
% Final passive clauses: 306
% Main loop iterations started: 205
% 
% >>> Simplifying Inferences
% Duplicate literals: 64
% Fw subsumption resolutions: 35
% Fw demodulations: 39
% Bw demodulations: 61
% 
% >>> Deletion Inferences
% Simple tautologies: 64
% Equational tautologies: 3
% Forward subsumptions: 442
% Backward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 455
% Unit resulting resolution: 4
% Factoring: 6
% Forward superposition: 271
% Backward superposition: 213
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 67
% Split components: 134
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 178
% SAT solver unit clauses: 8
% SAT solver binary clauses: 67
% 
% Memory used [KB]: 6012
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
% Time measurement results:
% sat proof minimization: 0.001 s
% congruence closure: 0.001 s
% forward subsumption: 0.002 s
% forward subsumption resolution: 0.001 s
% forward demodulation: 0.001 s
% unit resulting resolution: 0.001 s
% superposition: 0.002 s
% other: 0.011 s ( own 0.002 s ) 

