i: [1,24,0,11,1,1,0,0] l98_xboole_1
a: [3,36,0,11,1,24]
s: 36
i: [1,28,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 28
i: [1,29,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 29
i: [1,30,0,7,1,0,0,1] commutativity_k5_xboole_0
s: 30
i: [1,31,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 31
i: [1,32,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 32
i: [1,25,0,9,1,0,0,1] t47_xboole_1
s: 25
i: [1,26,0,11,1,0,0,1] d6_xboole_0
s: 26
d: [2,42,1,7,1,0,44,26,29]
s: 42
i: [1,27,0,13,1,0,0,1] t42_xboole_1
s: 27
d: [2,45,1,7,1,0,44,29,26]
s: 45
d: [2,38,1,9,1,0,44,25,31]
s: 38
d: [2,39,1,9,1,0,44,25,31]
s: 39
d: [2,43,1,11,1,0,44,26,32]
s: 43
d: [2,44,1,11,1,0,44,26,32]
s: 44
d: [2,46,1,11,1,0,44,32,26]
s: 46
d: [2,47,1,11,1,0,44,32,26]
s: 47
d: [2,57,1,13,1,0,44,27,32]
s: 57
d: [2,100,2,11,1,0,44,57,27]
s: 100
d: [2,58,1,13,1,0,44,27,32]
s: 58
d: [2,61,1,13,1,0,44,32,27]
s: 61
d: [2,106,2,11,1,0,44,27,57]
s: 106
d: [2,62,1,13,1,0,44,32,27]
s: 62
d: [2,115,3,15,1,0,44,100,26]
d: [2,150,3,11,1,0,22,115,26]
s: 150
d: [2,40,1,15,1,0,44,26,25]
s: 40
d: [2,41,1,15,1,0,44,26,25]
d: [2,48,1,15,1,0,22,41,30]
s: 48
d: [2,122,3,15,1,0,44,100,26]
d: [2,152,3,11,1,0,22,122,26]
s: 152
d: [2,51,1,17,1,0,44,27,25]
s: 51
d: [2,250,2,15,1,0,44,51,38]
d: [2,259,2,11,1,0,22,250,26]
s: 259
d: [2,276,3,7,1,1,44,36,259]
d: [2,291,3,0,0,1,19,276]
e: 291
d: [2,260,3,0,0,1,52,36,259]
e: 260
f: 260,24
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__l98_xboole_1
% SZS output start Proof for xboole_1__l98_xboole_1
1. ! [X0] : ! [X1] : k5_xboole_0(X0,X1) = k4_xboole_0(k2_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [input l98_xboole_1]
2. ~! [X0] : ! [X1] : k5_xboole_0(X0,X1) = k4_xboole_0(k2_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [negated conjecture 1]
4. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
6. ! [X0] : ! [X1] : k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [input d6_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : k4_xboole_0(k2_xboole_0(X0,X1),X2) = k2_xboole_0(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2)) [input t42_xboole_1]
14. ! [X0] : ! [X1] : k4_xboole_0(X0,X1) = k4_xboole_0(X0,k3_xboole_0(X0,X1)) [input t47_xboole_1]
15. ~! [X0,X1] : k5_xboole_0(X0,X1) = k4_xboole_0(k2_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [flattening 2]
16. ! [X0,X1] : k4_xboole_0(X0,X1) = k4_xboole_0(X0,k3_xboole_0(X0,X1)) [flattening 14]
17. ! [X0,X1] : k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [flattening 6]
18. ! [X0,X1,X2] : k4_xboole_0(k2_xboole_0(X0,X1),X2) = k2_xboole_0(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2)) [flattening 13]
21. ? [X0,X1] : k5_xboole_0(X0,X1) != k4_xboole_0(k2_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [ennf transformation 15]
22. ? [X0,X1] : k5_xboole_0(X0,X1) != k4_xboole_0(k2_xboole_0(X0,X1),k3_xboole_0(X0,X1)) => k5_xboole_0(sK0,sK1) != k4_xboole_0(k2_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) [choice axiom]
23. k5_xboole_0(sK0,sK1) != k4_xboole_0(k2_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) [skolemisation 21,22]
24. k5_xboole_0(sK0,sK1) != k4_xboole_0(k2_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) [cnf transformation 23] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
25. k4_xboole_0(X0,X1) = k4_xboole_0(X0,k3_xboole_0(X0,X1)) [cnf transformation 16] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
26. k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [cnf transformation 17] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
27. k4_xboole_0(k2_xboole_0(X0,X1),X2) = k2_xboole_0(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2)) [cnf transformation 18] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
31. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. 1 <=> k5_xboole_0(sK0,sK1) = k4_xboole_0(k2_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) [avatar definition]
36. k5_xboole_0(sK0,sK1) != k4_xboole_0(k2_xboole_0(sK0,sK1),k3_xboole_0(sK0,sK1)) <- (~1) [avatar component clause 34] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
37. ~1 [avatar split clause 24,34]
38. k4_xboole_0(X1,X2) = k4_xboole_0(X1,k3_xboole_0(X2,X1)) [superposition 25,31] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
51. k4_xboole_0(k2_xboole_0(X0,X2),k3_xboole_0(X0,X1)) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X2,k3_xboole_0(X0,X1))) [superposition 27,25] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
250. k4_xboole_0(k2_xboole_0(X3,X2),k3_xboole_0(X3,X2)) = k2_xboole_0(k4_xboole_0(X3,X2),k4_xboole_0(X2,X3)) [superposition 51,38] {a:2,w:15,wCS:6000,sine:1,thAx:0,allAx:4,thDist:-4}
259. k5_xboole_0(X3,X2) = k4_xboole_0(k2_xboole_0(X3,X2),k3_xboole_0(X3,X2)) [forward demodulation 250,26] {a:2,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
260. $false <- (~1) [unit resulting resolution 36,259] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
293. 1 [avatar contradiction clause 260]
303. $false [avatar sat refutation 37,293] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:15,thDist:-15}
% SZS output end Proof for xboole_1__l98_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 9
% Generated clauses: 273
% Active clauses: 19
% Passive clauses: 134
% Final active clauses: 19
% Final passive clauses: 105
% Main loop iterations started: 29
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw demodulations: 52
% 
% >>> Deletion Inferences
% Forward subsumptions: 56
% Fw demodulations to eq. taut.: 10
% 
% >>> Generating Inferences
% Unit resulting resolution: 1
% Forward superposition: 118
% Backward superposition: 91
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 5
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward demodulation: 0.003 s
% other: 0.003 s ( own 0.0000 s ) 

