i: [1,22,0,10,2,1,0,0] t1_gate_1
a: [3,34,0,5,1,22]
s: 34
i: [1,23,0,2,1,0,0,1] fc1_xboole_0
a: [3,43,0,2,1,23]
s: 43
i: [1,29,0,2,1,0,0,2] rc2_xboole_0
a: [3,48,0,2,1,29]
s: 48
i: [1,30,0,2,1,0,0,2] rc1_xboole_0
a: [3,53,0,2,1,30]
s: 53
i: [1,24,0,3,1,0,0,1] fc2_xboole_0
s: 24
i: [1,27,0,5,2,0,0,1] fc2_gate_1
s: 27
i: [1,28,0,5,2,0,0,1] fc1_gate_1
s: 28
i: [1,25,0,6,2,0,0,1] d1_gate_1
s: 25
d: [2,56,1,0,0,1,52,24,34,25]
e: 56
d: [2,57,1,6,2,1,44,34,25]
d: [2,61,1,3,1,1,19,57]
d: [2,63,1,0,0,1,21,61,24]
e: 63
r: 34
a: [3,33,0,5,1,22]
s: 33
a: [3,38,0,5,1,22]
s: 38
i: [1,26,0,7,2,0,0,1] d1_gate_1
s: 26
d: [2,68,1,0,0,2,52,43,38,26]
e: 68
d: [2,75,1,7,2,1,44,38,26]
d: [2,76,1,2,1,1,19,75]
d: [2,79,1,0,0,2,21,76,43]
e: 79
f: 79,63,23,22
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t1_gate_1
% SZS output start Proof for gate_1__t1_gate_1
1. k1_tarski(k1_xboole_0) = k1_gate_1(k1_xboole_0) & k1_xboole_0 = k1_gate_1(k1_tarski(k1_xboole_0)) [input t1_gate_1]
2. ~(k1_tarski(k1_xboole_0) = k1_gate_1(k1_xboole_0) & k1_xboole_0 = k1_gate_1(k1_tarski(k1_xboole_0))) [negated conjecture 1]
8. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
10. ! [X0] : ~v1_xboole_0(k1_tarski(X0)) [input fc2_xboole_0]
13. ! [X0] : ((v1_xboole_0(X0) => k1_tarski(k1_xboole_0) = k1_gate_1(X0)) & (~v1_xboole_0(X0) => k1_xboole_0 = k1_gate_1(X0))) [input d1_gate_1]
14. k1_tarski(k1_xboole_0) != k1_gate_1(k1_xboole_0) | k1_xboole_0 != k1_gate_1(k1_tarski(k1_xboole_0)) [ennf transformation 2]
15. ! [X0] : ((k1_tarski(k1_xboole_0) = k1_gate_1(X0) | ~v1_xboole_0(X0)) & (k1_xboole_0 = k1_gate_1(X0) | v1_xboole_0(X0))) [ennf transformation 13]
22. k1_tarski(k1_xboole_0) != k1_gate_1(k1_xboole_0) | k1_xboole_0 != k1_gate_1(k1_tarski(k1_xboole_0)) [cnf transformation 14] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
23. v1_xboole_0(k1_xboole_0) [cnf transformation 8] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
24. ~v1_xboole_0(k1_tarski(X0)) [cnf transformation 10] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
25. k1_xboole_0 = k1_gate_1(X0) | v1_xboole_0(X0) [cnf transformation 15] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
26. k1_tarski(k1_xboole_0) = k1_gate_1(X0) | ~v1_xboole_0(X0) [cnf transformation 15] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
32. 1 <=> k1_xboole_0 = k1_gate_1(k1_tarski(k1_xboole_0)) [avatar definition]
34. k1_xboole_0 != k1_gate_1(k1_tarski(k1_xboole_0)) <- (~1) [avatar component clause 32] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
36. 2 <=> k1_tarski(k1_xboole_0) = k1_gate_1(k1_xboole_0) [avatar definition]
38. k1_tarski(k1_xboole_0) != k1_gate_1(k1_xboole_0) <- (~2) [avatar component clause 36] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. ~1 | ~2 [avatar split clause 22,36,32]
41. 3 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
43. v1_xboole_0(k1_xboole_0) <- (3) [avatar component clause 41] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. 3 [avatar split clause 23,41]
57. k1_xboole_0 != k1_xboole_0 | v1_xboole_0(k1_tarski(k1_xboole_0)) <- (~1) [superposition 34,25] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
61. v1_xboole_0(k1_tarski(k1_xboole_0)) <- (~1) [trivial inequality removal 57] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
63. $false <- (~1) [subsumption resolution 61,24] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
64. 1 [avatar contradiction clause 63]
75. k1_tarski(k1_xboole_0) != k1_tarski(k1_xboole_0) | ~v1_xboole_0(k1_xboole_0) <- (~2) [superposition 38,26] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
76. ~v1_xboole_0(k1_xboole_0) <- (~2) [trivial inequality removal 75] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
79. $false <- (~2, 3) [subsumption resolution 76,43] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
80. 2 | ~3 [avatar contradiction clause 79]
81. $false [avatar sat refutation 39,44,64,80] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:21,thDist:-21}
% SZS output end Proof for gate_1__t1_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 9
% Generated clauses: 37
% Active clauses: 11
% Passive clauses: 11
% Final active clauses: 10
% Main loop iterations started: 11
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Trivial inequalities: 2
% Fw subsumption resolutions: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 1
% Unit resulting resolution: 2
% Forward superposition: 3
% Backward superposition: 10
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 5
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

