i: [1,46,0,2,1,1,0,0] t37_pboole
a: [3,83,0,2,1,46]
s: 83
i: [1,44,0,2,1,1,0,0] t37_pboole
a: [3,93,0,2,1,44]
s: 93
i: [1,42,0,2,1,1,0,0] t37_pboole
a: [3,103,0,2,1,42]
s: 103
i: [1,40,0,2,1,1,0,0] t37_pboole
a: [3,113,0,2,1,40]
s: 113
i: [1,38,0,2,1,1,0,0] t37_pboole
a: [3,123,0,2,1,38]
s: 123
i: [1,36,0,2,1,1,0,0] t37_pboole
a: [3,133,0,2,1,36]
s: 133
i: [1,47,0,3,1,1,0,0] t37_pboole
a: [3,78,0,3,1,47]
s: 78
i: [1,45,0,3,1,1,0,0] t37_pboole
a: [3,88,0,3,1,45]
s: 88
i: [1,43,0,3,1,1,0,0] t37_pboole
a: [3,98,0,3,1,43]
s: 98
i: [1,41,0,3,1,1,0,0] t37_pboole
a: [3,108,0,3,1,41]
s: 108
i: [1,39,0,3,1,1,0,0] t37_pboole
a: [3,118,0,3,1,39]
s: 118
i: [1,37,0,3,1,1,0,0] t37_pboole
a: [3,128,0,3,1,37]
s: 128
i: [1,49,0,4,1,1,0,0] t37_pboole
a: [3,68,0,4,1,49]
s: 68
i: [1,48,0,7,1,1,0,0] t37_pboole
a: [3,73,0,7,1,48]
s: 73
i: [1,64,0,2,1,0,0,1] rc1_funct_1
a: [3,138,0,2,1,64]
s: 138
i: [1,63,0,2,1,0,0,1] rc1_funct_1
a: [3,143,0,2,1,63]
s: 143
i: [1,52,0,3,1,0,0,1] rc1_pboole
s: 52
i: [1,54,0,3,1,0,0,1] rc1_pboole
s: 54
i: [1,53,0,4,1,0,0,1] rc1_pboole
s: 53
i: [1,55,0,4,1,0,0,1] rc1_pboole
s: 55
i: [1,56,0,24,9,0,0,1] reflexivity_r2_pboole
s: 56
i: [1,59,0,25,9,0,0,1] dt_k2_pboole
s: 59
i: [1,61,0,25,9,0,0,1] dt_k2_pboole
s: 61
i: [1,57,0,26,9,0,0,255] idempotence_k2_pboole
s: 57
i: [1,60,0,26,9,0,0,1] dt_k2_pboole
s: 60
i: [1,62,0,26,9,0,0,1] dt_k2_pboole
s: 62
i: [1,50,0,27,9,0,0,1] t14_pboole
s: 50
i: [1,58,0,29,9,0,0,255] commutativity_k2_pboole
s: 58
d: [2,160,1,47,17,0,44,50,58]
d: [2,169,1,27,9,0,18,160]
s: 169
i: [1,51,0,42,15,0,0,1] t13_pboole
s: 51
d: [2,181,1,53,14,1,40,51,73]
d: [2,188,1,50,13,2,21,181,78]
d: [2,189,1,48,12,3,21,188,83]
d: [2,190,1,45,11,4,21,189,88]
d: [2,191,1,43,10,5,21,190,93]
a: [3,195,1,5,1,191]
s: 195
d: [2,220,2,0,0,9,52,133,113,123,103,128,118,108,98,195,59]
e: 220
d: [2,221,2,20,8,1,40,195,59]
d: [2,223,2,17,7,2,21,221,98]
d: [2,224,2,15,6,3,21,223,103]
d: [2,225,2,12,5,4,21,224,108]
d: [2,226,2,10,4,5,21,225,113]
d: [2,227,2,7,3,6,21,226,118]
d: [2,228,2,5,2,7,21,227,123]
d: [2,229,2,2,1,8,21,228,128]
d: [2,230,2,0,0,9,21,229,133]
e: 230
r: 195
a: [3,194,1,5,1,191]
s: 194
a: [3,199,1,6,1,191]
s: 199
d: [2,232,2,0,0,9,52,133,113,123,103,128,118,108,98,199,60]
e: 232
d: [2,233,2,20,8,1,40,199,60]
d: [2,235,2,17,7,2,21,233,98]
d: [2,236,2,15,6,3,21,235,103]
d: [2,237,2,12,5,4,21,236,108]
d: [2,238,2,10,4,5,21,237,113]
d: [2,239,2,7,3,6,21,238,118]
d: [2,240,2,5,2,7,21,239,123]
d: [2,241,2,2,1,8,21,240,128]
d: [2,242,2,0,0,9,21,241,133]
e: 242
r: 199
a: [3,203,1,5,1,191]
s: 203
d: [2,244,2,0,0,9,52,133,113,123,103,128,118,108,98,203,61]
e: 244
d: [2,245,2,20,8,1,40,203,61]
d: [2,247,2,17,7,2,21,245,98]
d: [2,248,2,15,6,3,21,247,103]
d: [2,249,2,12,5,4,21,248,108]
d: [2,250,2,10,4,5,21,249,113]
d: [2,251,2,7,3,6,21,250,118]
d: [2,252,2,5,2,7,21,251,123]
d: [2,253,2,2,1,8,21,252,128]
d: [2,254,2,0,0,9,21,253,133]
e: 254
r: 203
a: [3,202,1,5,1,191]
s: 202
a: [3,198,1,6,1,191]
s: 198
a: [3,207,1,6,1,191]
s: 207
d: [2,256,2,0,0,9,52,133,113,123,103,128,118,108,98,207,62]
e: 256
d: [2,257,2,20,8,1,40,207,62]
d: [2,259,2,17,7,2,21,257,98]
d: [2,260,2,15,6,3,21,259,103]
d: [2,261,2,12,5,4,21,260,108]
d: [2,262,2,10,4,5,21,261,113]
d: [2,263,2,7,3,6,21,262,118]
d: [2,264,2,5,2,7,21,263,123]
d: [2,265,2,2,1,8,21,264,128]
d: [2,266,2,0,0,9,21,265,133]
e: 266
r: 207
a: [3,206,1,6,1,191]
s: 206
a: [3,210,1,21,6,191]
s: 210
d: [2,268,2,34,13,1,40,210,50]
d: [2,273,2,24,9,1,18,268]
d: [2,274,2,22,8,2,21,273,133]
d: [2,275,2,19,7,3,21,274,128]
d: [2,276,2,17,6,4,21,275,123]
d: [2,277,2,14,5,5,21,276,118]
d: [2,278,2,10,4,6,21,277,68]
d: [2,279,2,7,3,7,21,278,98]
d: [2,280,2,5,2,8,21,279,103]
d: [2,281,2,2,1,9,21,280,108]
d: [2,282,2,0,0,10,21,281,113]
e: 282
f: 282,266,254,242,230,191,36,37,38,39,40,41,42,43,44,45,46,47,48,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pboole__t37_pboole
% SZS output start Proof for pboole__t37_pboole
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => (r2_pboole(X0,k2_pboole(X0,X1,X2),X3) => r2_pboole(X0,X1,X3))))) [input t37_pboole]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => (r2_pboole(X0,k2_pboole(X0,X1,X2),X3) => r2_pboole(X0,X1,X3))))) [negated conjecture 1]
4. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2)))) [input dt_k2_pboole]
9. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,X3))))) [input t13_pboole]
10. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,X1,k2_pboole(X0,X1,X2)))) [input t14_pboole]
11. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => (r2_pboole(X0,k2_pboole(X0,X1,X2),X3) => r2_pboole(X0,X1,X3))))) [flattening 2]
12. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,X1,k2_pboole(X0,X1,X2)))) [flattening 10]
13. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,X3))))) [flattening 9]
14. ? [X0,X1] : (? [X2] : (? [X3] : ((~r2_pboole(X0,X1,X3) & r2_pboole(X0,k2_pboole(X0,X1,X2),X3)) & (v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3))) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 11]
15. ? [X0,X1] : (? [X2] : (? [X3] : (~r2_pboole(X0,X1,X3) & r2_pboole(X0,k2_pboole(X0,X1,X2),X3) & v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 14]
16. ! [X0,X1] : (! [X2] : (r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 12]
17. ! [X0,X1] : (! [X2] : (r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 16]
18. ! [X0,X1] : (! [X2] : (! [X3] : ((r2_pboole(X0,X1,X3) | (~r2_pboole(X0,X2,X3) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 13]
19. ! [X0,X1] : (! [X2] : (! [X3] : (r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X2,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 18]
26. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 4]
27. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 26]
28. ? [X0,X1] : (? [X2] : (? [X3] : (~r2_pboole(X0,X1,X3) & r2_pboole(X0,k2_pboole(X0,X1,X2),X3) & v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (~r2_pboole(sK0,sK1,X3) & r2_pboole(sK0,k2_pboole(sK0,sK1,X2),X3) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
29. ? [X2] : (? [X3] : (~r2_pboole(sK0,sK1,X3) & r2_pboole(sK0,k2_pboole(sK0,sK1,X2),X3) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (? [X3] : (~r2_pboole(sK0,sK1,X3) & r2_pboole(sK0,k2_pboole(sK0,sK1,sK2),X3) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
30. ? [X3] : (~r2_pboole(sK0,sK1,X3) & r2_pboole(sK0,k2_pboole(sK0,sK1,sK2),X3) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) => (~r2_pboole(sK0,sK1,sK3) & r2_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK3) & v1_partfun1(sK3,sK0) & v1_funct_1(sK3) & v4_relat_1(sK3,sK0) & v1_relat_1(sK3)) [choice axiom]
31. ((~r2_pboole(sK0,sK1,sK3) & r2_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK3) & v1_partfun1(sK3,sK0) & v1_funct_1(sK3) & v4_relat_1(sK3,sK0) & v1_relat_1(sK3)) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 15,30,29,28]
36. v1_relat_1(sK1) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
37. v4_relat_1(sK1,sK0) [cnf transformation 31] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
38. v1_funct_1(sK1) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
39. v1_partfun1(sK1,sK0) [cnf transformation 31] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
40. v1_relat_1(sK2) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
41. v4_relat_1(sK2,sK0) [cnf transformation 31] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
42. v1_funct_1(sK2) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
43. v1_partfun1(sK2,sK0) [cnf transformation 31] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
44. v1_relat_1(sK3) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
45. v4_relat_1(sK3,sK0) [cnf transformation 31] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
46. v1_funct_1(sK3) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
47. v1_partfun1(sK3,sK0) [cnf transformation 31] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
48. r2_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK3) [cnf transformation 31] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
49. ~r2_pboole(sK0,sK1,sK3) [cnf transformation 31] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
50. r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 17] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. ~r2_pboole(X0,X2,X3) | r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 19] {a:0,w:42,wCS:16800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. v1_relat_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 27] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. v4_relat_1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 27] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. v1_funct_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 27] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. v1_partfun1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 27] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
66. 1 <=> r2_pboole(sK0,sK1,sK3) [avatar definition]
68. ~r2_pboole(sK0,sK1,sK3) <- (~1) [avatar component clause 66] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
69. ~1 [avatar split clause 49,66]
71. 2 <=> r2_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK3) [avatar definition]
73. r2_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK3) <- (2) [avatar component clause 71] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
74. 2 [avatar split clause 48,71]
76. 3 <=> v1_partfun1(sK3,sK0) [avatar definition]
78. v1_partfun1(sK3,sK0) <- (3) [avatar component clause 76] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
79. 3 [avatar split clause 47,76]
81. 4 <=> v1_funct_1(sK3) [avatar definition]
83. v1_funct_1(sK3) <- (4) [avatar component clause 81] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
84. 4 [avatar split clause 46,81]
86. 5 <=> v4_relat_1(sK3,sK0) [avatar definition]
88. v4_relat_1(sK3,sK0) <- (5) [avatar component clause 86] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
89. 5 [avatar split clause 45,86]
91. 6 <=> v1_relat_1(sK3) [avatar definition]
93. v1_relat_1(sK3) <- (6) [avatar component clause 91] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
94. 6 [avatar split clause 44,91]
96. 7 <=> v1_partfun1(sK2,sK0) [avatar definition]
98. v1_partfun1(sK2,sK0) <- (7) [avatar component clause 96] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
99. 7 [avatar split clause 43,96]
101. 8 <=> v1_funct_1(sK2) [avatar definition]
103. v1_funct_1(sK2) <- (8) [avatar component clause 101] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
104. 8 [avatar split clause 42,101]
106. 9 <=> v4_relat_1(sK2,sK0) [avatar definition]
108. v4_relat_1(sK2,sK0) <- (9) [avatar component clause 106] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
109. 9 [avatar split clause 41,106]
111. 10 <=> v1_relat_1(sK2) [avatar definition]
113. v1_relat_1(sK2) <- (10) [avatar component clause 111] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
114. 10 [avatar split clause 40,111]
116. 11 <=> v1_partfun1(sK1,sK0) [avatar definition]
118. v1_partfun1(sK1,sK0) <- (11) [avatar component clause 116] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
119. 11 [avatar split clause 39,116]
121. 12 <=> v1_funct_1(sK1) [avatar definition]
123. v1_funct_1(sK1) <- (12) [avatar component clause 121] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
124. 12 [avatar split clause 38,121]
126. 13 <=> v4_relat_1(sK1,sK0) [avatar definition]
128. v4_relat_1(sK1,sK0) <- (13) [avatar component clause 126] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
129. 13 [avatar split clause 37,126]
131. 14 <=> v1_relat_1(sK1) [avatar definition]
133. v1_relat_1(sK1) <- (14) [avatar component clause 131] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
134. 14 [avatar split clause 36,131]
181. r2_pboole(sK0,X0,sK3) | ~r2_pboole(sK0,X0,k2_pboole(sK0,sK1,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (2) [resolution 51,73] {a:1,w:53,wCS:5300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
188. r2_pboole(sK0,X0,sK3) | ~r2_pboole(sK0,X0,k2_pboole(sK0,sK1,sK2)) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (2, 3) [subsumption resolution 181,78] {a:1,w:50,wCS:5000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
189. r2_pboole(sK0,X0,sK3) | ~r2_pboole(sK0,X0,k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (2, 3, 4) [subsumption resolution 188,83] {a:1,w:48,wCS:4800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
190. r2_pboole(sK0,X0,sK3) | ~r2_pboole(sK0,X0,k2_pboole(sK0,sK1,sK2)) | ~v1_relat_1(sK3) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (2, 3, 4, 5) [subsumption resolution 189,88] {a:1,w:45,wCS:4500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
191. r2_pboole(sK0,X0,sK3) | ~r2_pboole(sK0,X0,k2_pboole(sK0,sK1,sK2)) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (2, 3, 4, 5, 6) [subsumption resolution 190,93] {a:1,w:43,wCS:4300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
193. 17 <=> v1_relat_1(k2_pboole(sK0,sK1,sK2)) [avatar definition]
195. ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (~17) [avatar component clause 193] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
197. 18 <=> v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) [avatar definition]
199. ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) <- (~18) [avatar component clause 197] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
201. 19 <=> v1_funct_1(k2_pboole(sK0,sK1,sK2)) [avatar definition]
203. ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) <- (~19) [avatar component clause 201] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
205. 20 <=> v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) [avatar definition]
207. ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) <- (~20) [avatar component clause 205] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
209. 21 <=> ! [X0] : (r2_pboole(sK0,X0,sK3) | ~v1_relat_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_funct_1(X0) | ~v1_partfun1(X0,sK0) | ~r2_pboole(sK0,X0,k2_pboole(sK0,sK1,sK2))) [avatar definition]
210. ~r2_pboole(sK0,X0,k2_pboole(sK0,sK1,sK2)) | ~v1_relat_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_funct_1(X0) | ~v1_partfun1(X0,sK0) | r2_pboole(sK0,X0,sK3) <- (21) [avatar component clause 209] {a:1,w:21,wCS:2100,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
211. ~17 | ~18 | ~19 | ~20 | 21 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar split clause 191,91,86,81,76,71,209,205,201,197,193]
221. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~17) [resolution 195,59] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
223. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, ~17) [subsumption resolution 221,98] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
224. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, ~17) [subsumption resolution 223,103] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
225. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, ~17) [subsumption resolution 224,108] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
226. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, ~17) [subsumption resolution 225,113] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
227. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, ~17) [subsumption resolution 226,118] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
228. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, 12, ~17) [subsumption resolution 227,123] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
229. ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, 12, 13, ~17) [subsumption resolution 228,128] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
230. $false <- (7, 8, 9, 10, 11, 12, 13, 14, ~17) [subsumption resolution 229,133] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
231. ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | 17 [avatar contradiction clause 230]
233. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~18) [resolution 199,60] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
235. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, ~18) [subsumption resolution 233,98] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
236. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, ~18) [subsumption resolution 235,103] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
237. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, ~18) [subsumption resolution 236,108] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
238. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, ~18) [subsumption resolution 237,113] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
239. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, ~18) [subsumption resolution 238,118] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
240. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, 12, ~18) [subsumption resolution 239,123] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
241. ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, 12, 13, ~18) [subsumption resolution 240,128] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
242. $false <- (7, 8, 9, 10, 11, 12, 13, 14, ~18) [subsumption resolution 241,133] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
243. ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | 18 [avatar contradiction clause 242]
245. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~19) [resolution 203,61] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
247. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, ~19) [subsumption resolution 245,98] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
248. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, ~19) [subsumption resolution 247,103] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
249. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, ~19) [subsumption resolution 248,108] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
250. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, ~19) [subsumption resolution 249,113] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
251. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, ~19) [subsumption resolution 250,118] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
252. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, 12, ~19) [subsumption resolution 251,123] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
253. ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, 12, 13, ~19) [subsumption resolution 252,128] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
254. $false <- (7, 8, 9, 10, 11, 12, 13, 14, ~19) [subsumption resolution 253,133] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
255. ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | 19 [avatar contradiction clause 254]
257. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~20) [resolution 207,62] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
259. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, ~20) [subsumption resolution 257,98] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
260. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, ~20) [subsumption resolution 259,103] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
261. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, ~20) [subsumption resolution 260,108] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
262. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, ~20) [subsumption resolution 261,113] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
263. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, ~20) [subsumption resolution 262,118] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
264. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, 12, ~20) [subsumption resolution 263,123] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
265. ~v1_relat_1(sK1) <- (7, 8, 9, 10, 11, 12, 13, ~20) [subsumption resolution 264,128] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
266. $false <- (7, 8, 9, 10, 11, 12, 13, 14, ~20) [subsumption resolution 265,133] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
267. ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | 20 [avatar contradiction clause 266]
268. ~v1_relat_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_funct_1(sK1) | ~v1_partfun1(sK1,sK0) | r2_pboole(sK0,sK1,sK3) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (21) [resolution 210,50] {a:2,w:34,wCS:3400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
273. ~v1_relat_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_funct_1(sK1) | ~v1_partfun1(sK1,sK0) | r2_pboole(sK0,sK1,sK3) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (21) [duplicate literal removal 268] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
274. ~v4_relat_1(sK1,sK0) | ~v1_funct_1(sK1) | ~v1_partfun1(sK1,sK0) | r2_pboole(sK0,sK1,sK3) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (14, 21) [subsumption resolution 273,133] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
275. ~v1_funct_1(sK1) | ~v1_partfun1(sK1,sK0) | r2_pboole(sK0,sK1,sK3) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (13, 14, 21) [subsumption resolution 274,128] {a:2,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
276. ~v1_partfun1(sK1,sK0) | r2_pboole(sK0,sK1,sK3) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (12, 13, 14, 21) [subsumption resolution 275,123] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
277. r2_pboole(sK0,sK1,sK3) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (11, 12, 13, 14, 21) [subsumption resolution 276,118] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
278. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 11, 12, 13, 14, 21) [subsumption resolution 277,68] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
279. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 7, 11, 12, 13, 14, 21) [subsumption resolution 278,98] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
280. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1, 7, 8, 11, 12, 13, 14, 21) [subsumption resolution 279,103] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
281. ~v1_relat_1(sK2) <- (~1, 7, 8, 9, 11, 12, 13, 14, 21) [subsumption resolution 280,108] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
282. $false <- (~1, 7, 8, 9, 10, 11, 12, 13, 14, 21) [subsumption resolution 281,113] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
283. 1 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | ~21 [avatar contradiction clause 282]
292. $false [avatar sat refutation 69,74,79,84,89,94,99,104,109,114,119,124,129,134,211,231,243,255,267,283] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:273,thDist:-273}
% SZS output end Proof for pboole__t37_pboole
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 172
% Active clauses: 39
% Passive clauses: 42
% Final active clauses: 35
% Final passive clauses: 3
% Main loop iterations started: 39
% 
% >>> Simplifying Inferences
% Duplicate literals: 172
% Fw subsumption resolutions: 61
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 4
% Forward superposition: 6
% Backward superposition: 12
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 5
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 28
% SAT solver unit clauses: 16
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

