i: [1,172,0,10,3,1,0,0] t34_yellow19
a: [3,239,0,2,1,172]
s: 239
i: [1,164,0,2,1,1,0,0] t34_yellow19
a: [3,280,0,2,1,164]
s: 280
i: [1,163,0,2,1,1,0,0] t34_yellow19
a: [3,285,0,2,1,163]
s: 285
i: [1,162,0,2,1,1,0,0] t34_yellow19
a: [3,290,0,2,1,162]
s: 290
i: [1,173,0,2,1,0,0,255] spc1_boole
a: [3,295,0,2,1,173]
s: 295
i: [1,174,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,300,0,2,1,174]
s: 300
i: [1,225,0,2,1,0,0,255] existence_l1_orders_2
a: [3,310,0,2,1,225]
s: 310
i: [1,226,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,315,0,2,1,226]
s: 315
i: [1,227,0,2,1,0,0,2] existence_l1_struct_0
a: [3,320,0,2,1,227]
s: 320
i: [1,229,0,2,1,0,0,255] rc1_relat_1
a: [3,325,0,2,1,229]
s: 325
i: [1,228,0,2,1,0,0,255] rc1_relat_1
a: [3,330,0,2,1,228]
s: 330
i: [1,231,0,2,1,0,0,255] rc1_finset_1
a: [3,335,0,2,1,231]
s: 335
i: [1,230,0,2,1,0,0,255] rc1_finset_1
a: [3,340,0,2,1,230]
s: 340
i: [1,234,0,2,1,0,0,255] rc8_finset_1
a: [3,345,0,2,1,234]
s: 345
i: [1,233,0,2,1,0,0,255] rc8_finset_1
a: [3,350,0,2,1,233]
s: 350
i: [1,232,0,2,1,0,0,255] rc8_finset_1
a: [3,355,0,2,1,232]
s: 355
i: [1,235,0,2,1,0,0,255] rc2_relat_1
a: [3,360,0,2,1,235]
s: 360
i: [1,175,0,3,1,0,0,255] d2_xboole_0
a: [3,305,0,3,1,175]
s: 305
i: [1,176,0,4,2,0,0,255] cc5_finset_1
s: 176
i: [1,177,0,4,2,0,0,255] cc4_finset_1
s: 177
i: [1,178,0,4,2,0,0,255] dt_l1_orders_2
s: 178
d: [2,362,1,2,1,1,40,178,310]
a: [3,366,1,2,1,362]
s: 366
i: [1,180,0,4,2,0,0,1] dt_l1_pre_topc
s: 180
d: [2,368,1,2,1,1,40,180,280]
a: [3,373,1,2,1,368]
s: 373
i: [1,181,0,4,2,0,0,255] cc6_finset_1
s: 181
i: [1,182,0,4,2,0,0,255] cc1_relat_1
s: 182
d: [2,369,1,2,1,1,40,180,315]
a: [3,378,1,2,1,369]
s: 378
i: [1,183,0,4,2,0,0,255] cc1_finset_1
s: 183
d: [2,380,1,2,1,1,40,181,300]
a: [3,384,1,2,1,380]
s: 384
i: [1,223,0,4,1,0,0,1] existence_m1_subset_1
s: 223
i: [1,166,0,20,7,1,0,0] t34_yellow19
a: [3,271,0,18,6,166]
s: 271
d: [2,386,1,2,1,1,40,182,300]
a: [3,390,1,2,1,386]
s: 390
i: [1,165,0,20,7,1,0,0] t34_yellow19
a: [3,275,0,18,6,165]
s: 275
d: [2,392,1,2,1,1,40,183,300]
a: [3,396,1,2,1,392]
s: 396
i: [1,184,0,5,2,0,0,255] t6_boole
s: 184
i: [1,222,0,5,2,0,0,1] t7_boole
s: 222
d: [2,398,0,5,2,1,22,184,305]
s: 398
i: [1,185,0,6,3,0,0,255] cc7_struct_0
s: 185
i: [1,186,0,6,3,0,0,255] cc5_struct_0
s: 186
i: [1,188,0,6,3,0,0,255] cc17_struct_0
s: 188
i: [1,189,0,6,3,0,0,255] cc2_struct_0
s: 189
d: [2,407,1,4,2,1,40,186,320]
a: [3,429,1,2,1,407]
s: 429
i: [1,190,0,6,3,0,0,255] cc6_struct_0
s: 190
d: [2,405,2,4,2,1,40,186,366]
a: [3,411,2,2,1,405]
s: 411
i: [1,191,0,6,3,0,0,255] cc1_struct_0
s: 191
i: [1,194,0,6,3,0,0,255] cc4_struct_0
s: 194
d: [2,406,2,4,2,1,40,186,378]
a: [3,420,2,2,1,406]
s: 420
i: [1,200,0,6,2,0,0,1] existence_l1_waybel_0
s: 200
i: [1,219,0,6,2,0,0,1] t1_subset
s: 219
i: [1,224,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 224
i: [1,179,0,7,3,0,0,255] cc7_finset_1
s: 179
d: [2,455,1,19,7,1,40,179,275]
a: [3,459,1,3,1,455]
s: 459
i: [1,192,0,7,3,0,0,255] cc8_struct_0
s: 192
d: [2,454,1,5,2,0,40,179,223]
s: 454
i: [1,195,0,7,3,0,0,255] cc11_struct_0
s: 195
i: [1,196,0,7,3,0,0,255] cc11_struct_0
s: 196
i: [1,197,0,7,3,0,0,255] cc9_struct_0
s: 197
i: [1,199,0,7,3,0,0,255] dt_l1_waybel_0
s: 199
d: [2,466,1,7,3,0,40,199,200]
d: [2,467,1,5,2,0,18,466]
s: 467
i: [1,201,0,7,3,0,0,255] fc9_struct_0
s: 201
i: [1,202,0,7,3,0,0,255] fc6_struct_0
s: 202
d: [2,468,2,5,2,0,40,467,178]
s: 468
i: [1,211,0,7,3,0,0,1] fc18_yellow_6
s: 211
i: [1,212,0,7,3,0,0,1] fc2_struct_0
s: 212
i: [1,216,0,7,3,0,0,255] fc8_struct_0
s: 216
i: [1,217,0,7,3,0,0,255] fc7_struct_0
s: 217
i: [1,218,0,7,3,0,0,1] fc1_struct_0
s: 218
i: [1,221,0,7,3,0,0,255] t8_boole
s: 221
d: [2,464,0,7,3,1,22,192,305]
s: 464
d: [2,465,0,7,3,1,22,197,305]
s: 465
i: [1,220,0,8,3,0,0,1] t2_subset
s: 220
d: [2,486,1,21,7,1,40,220,275]
a: [3,489,1,18,6,486]
s: 489
i: [1,198,0,9,4,0,0,255] cc10_struct_0
s: 198
a: [3,492,1,3,1,486]
s: 492
i: [1,205,0,11,5,0,0,1] l36_yellow19
s: 205
i: [1,206,0,11,5,0,0,1] l36_yellow19
s: 206
d: [2,495,2,18,6,1,40,489,224]
s: 495
i: [1,207,0,11,5,0,0,1] l36_yellow19
s: 207
d: [2,485,1,6,2,0,40,220,223]
s: 485
i: [1,208,0,12,5,0,0,1] l36_yellow19
s: 208
i: [1,209,0,13,5,0,0,1] l36_yellow19
s: 209
d: [2,507,2,6,2,0,40,485,224]
s: 507
i: [1,210,0,17,6,0,0,1] l36_yellow19
s: 210
d: [2,515,1,32,10,1,40,210,271]
d: [2,516,1,30,9,2,21,515,239]
d: [2,517,1,28,8,3,21,516,280]
d: [2,518,1,26,7,4,21,517,285]
d: [2,519,1,24,6,5,21,518,290]
d: [2,520,1,18,5,6,21,519,275]
a: [3,524,1,5,1,520]
s: 524
d: [2,542,2,0,0,5,52,290,285,280,239,524,209]
e: 542
d: [2,543,2,8,4,1,40,524,209]
d: [2,545,2,6,3,2,21,543,239]
d: [2,546,2,4,2,3,21,545,280]
d: [2,547,2,2,1,4,21,546,285]
d: [2,548,2,0,0,5,21,547,290]
e: 548
r: 524
i: [1,203,0,23,9,0,0,1] l35_yellow19
s: 203
i: [1,204,0,23,9,0,0,1] l35_yellow19
s: 204
a: [3,528,1,4,1,520]
s: 528
d: [2,554,2,0,0,5,52,290,285,280,239,528,208]
e: 554
d: [2,555,2,8,4,1,40,528,208]
d: [2,557,2,6,3,2,21,555,239]
d: [2,558,2,4,2,3,21,557,280]
d: [2,559,2,2,1,4,21,558,285]
d: [2,560,2,0,0,5,21,559,290]
e: 560
r: 528
a: [3,532,1,3,1,520]
s: 532
d: [2,562,2,0,0,5,52,290,285,280,239,532,207]
e: 562
d: [2,563,2,8,4,1,40,532,207]
d: [2,565,2,6,3,2,21,563,239]
d: [2,566,2,4,2,3,21,565,280]
d: [2,567,2,2,1,4,21,566,285]
d: [2,568,2,0,0,5,21,567,290]
e: 568
r: 532
a: [3,536,1,3,1,520]
s: 536
d: [2,570,2,0,0,5,52,290,285,280,239,536,206]
e: 570
d: [2,571,2,8,4,1,40,536,206]
d: [2,573,2,6,3,2,21,571,239]
d: [2,574,2,4,2,3,21,573,280]
d: [2,575,2,2,1,4,21,574,285]
d: [2,576,2,0,0,5,21,575,290]
e: 576
r: 536
a: [3,540,1,3,1,520]
s: 540
d: [2,578,2,0,0,5,52,290,285,280,239,540,205]
e: 578
d: [2,579,2,8,4,1,40,540,205]
d: [2,581,2,6,3,2,21,579,239]
d: [2,582,2,4,2,3,21,581,280]
d: [2,583,2,2,1,4,21,582,285]
d: [2,584,2,0,0,5,21,583,290]
e: 584
r: 239
r: 540
a: [3,238,0,2,1,172]
s: 238
i: [1,169,0,4,2,1,0,0] t34_yellow19
a: [3,257,0,2,1,169]
s: 257
i: [1,168,0,4,2,1,0,0] t34_yellow19
a: [3,262,0,2,1,168]
s: 262
i: [1,167,0,4,2,1,0,0] t34_yellow19
a: [3,267,0,2,1,167]
s: 267
i: [1,170,0,5,2,1,0,0] t34_yellow19
a: [3,252,0,3,1,170]
s: 252
i: [1,171,0,6,2,1,0,0] t34_yellow19
a: [3,247,0,4,1,171]
s: 247
d: [2,591,1,4,2,1,40,252,199]
d: [2,592,1,2,1,2,21,591,373]
a: [3,596,1,2,1,592]
s: 596
a: [3,242,0,8,2,172]
s: 242
d: [2,616,2,2,1,1,40,596,178]
a: [3,620,2,2,1,616]
s: 620
d: [2,600,1,3,1,1,40,247,222]
a: [3,614,1,3,1,600]
s: 614
d: [2,598,1,4,1,1,40,247,224]
a: [3,604,1,4,1,598]
s: 604
d: [2,599,1,4,1,1,40,247,219]
a: [3,609,1,4,1,599]
s: 609
d: [2,622,1,18,6,2,40,242,271]
d: [2,624,1,16,5,3,21,622,267]
d: [2,625,1,14,4,4,21,624,262]
d: [2,626,1,12,3,5,21,625,257]
d: [2,627,1,9,2,6,21,626,252]
d: [2,628,1,5,1,7,21,627,247]
a: [3,632,1,5,1,628]
s: 632
d: [2,660,2,0,0,7,52,267,262,257,252,247,632,275]
e: 660
d: [2,661,2,13,5,2,40,632,275]
d: [2,663,2,11,4,3,21,661,267]
d: [2,664,2,9,3,4,21,663,262]
d: [2,665,2,7,2,5,21,664,257]
d: [2,666,2,4,1,6,21,665,252]
d: [2,667,2,0,0,7,21,666,247]
e: 667
r: 275
r: 495
r: 489
d: [2,650,2,5,2,1,40,609,179]
a: [3,654,2,3,1,650]
s: 654
d: [2,623,1,23,9,1,40,242,204]
d: [2,634,1,20,8,2,21,623,252]
d: [2,635,1,18,7,3,21,634,257]
d: [2,636,1,16,6,4,21,635,262]
d: [2,637,1,14,5,5,21,636,267]
d: [2,638,1,12,4,6,21,637,238]
d: [2,639,1,10,3,7,21,638,280]
d: [2,640,1,8,2,8,21,639,285]
d: [2,641,1,6,1,9,21,640,290]
a: [3,645,1,6,1,641]
s: 645
d: [2,670,2,0,0,9,52,290,285,280,238,267,262,257,252,645,203]
e: 670
d: [2,671,2,17,8,1,40,645,203]
d: [2,673,2,14,7,2,21,671,252]
d: [2,674,2,12,6,3,21,673,257]
d: [2,675,2,10,5,4,21,674,262]
d: [2,676,2,8,4,5,21,675,267]
d: [2,677,2,6,3,6,21,676,238]
d: [2,678,2,4,2,7,21,677,280]
d: [2,679,2,2,1,8,21,678,285]
d: [2,680,2,0,0,9,21,679,290]
e: 680
d: [2,586,1,30,9,2,21,515,280]
d: [2,587,1,28,8,3,21,586,285]
d: [2,588,1,26,7,4,21,587,290]
d: [2,589,1,20,6,5,21,588,275]
f: 680,641,589,584,576,568,560,548,162,163,164,165,166,167,168,169,170,172
% Refutation found. Thanks to Tanya!
% SZS status Theorem for yellow19__t34_yellow19
% SZS output start Proof for yellow19__t34_yellow19
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (v1_compts_1(X0) <=> ! [X1] : ((l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) => ~(! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ~r3_waybel_9(X0,X1,X2)) & r2_hidden(X1,k6_yellow_6(X0)))))) [input t34_yellow19]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (v1_compts_1(X0) <=> ! [X1] : ((l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) => ~(! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ~r3_waybel_9(X0,X1,X2)) & r2_hidden(X1,k6_yellow_6(X0)))))) [negated conjecture 1]
57. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (v1_compts_1(X0) => ! [X1] : ((l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) => ? [X2] : (r3_waybel_9(X0,X1,X2) & m1_subset_1(X2,u1_struct_0(X0)))))) [input l35_yellow19]
58. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (! [X1] : ((l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) => ~(! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ~r3_waybel_9(X0,X1,X2)) & r2_hidden(X1,k6_yellow_6(X0)))) => v1_compts_1(X0))) [input l36_yellow19]
68. ? [X0] : ((v1_compts_1(X0) <~> ! [X1] : ((? [X2] : (r3_waybel_9(X0,X1,X2) & m1_subset_1(X2,u1_struct_0(X0))) | ~r2_hidden(X1,k6_yellow_6(X0))) | (~l1_waybel_0(X1,X0) | ~v7_waybel_0(X1) | ~v4_orders_2(X1) | v2_struct_0(X1)))) & (l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
69. ? [X0] : ((v1_compts_1(X0) <~> ! [X1] : (? [X2] : (r3_waybel_9(X0,X1,X2) & m1_subset_1(X2,u1_struct_0(X0))) | ~r2_hidden(X1,k6_yellow_6(X0)) | ~l1_waybel_0(X1,X0) | ~v7_waybel_0(X1) | ~v4_orders_2(X1) | v2_struct_0(X1))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [flattening 68]
107. ! [X0] : ((! [X1] : (? [X2] : (r3_waybel_9(X0,X1,X2) & m1_subset_1(X2,u1_struct_0(X0))) | (~l1_waybel_0(X1,X0) | ~v7_waybel_0(X1) | ~v4_orders_2(X1) | v2_struct_0(X1))) | ~v1_compts_1(X0)) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0))) [ennf transformation 57]
108. ! [X0] : (! [X1] : (? [X2] : (r3_waybel_9(X0,X1,X2) & m1_subset_1(X2,u1_struct_0(X0))) | ~l1_waybel_0(X1,X0) | ~v7_waybel_0(X1) | ~v4_orders_2(X1) | v2_struct_0(X1)) | ~v1_compts_1(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [flattening 107]
109. ! [X0] : ((v1_compts_1(X0) | ? [X1] : ((! [X2] : (~r3_waybel_9(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) & r2_hidden(X1,k6_yellow_6(X0))) & (l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0))) [ennf transformation 58]
110. ! [X0] : (v1_compts_1(X0) | ? [X1] : (! [X2] : (~r3_waybel_9(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) & r2_hidden(X1,k6_yellow_6(X0)) & l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [flattening 109]
133. ? [X0] : (((? [X1] : (! [X2] : (~r3_waybel_9(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) & r2_hidden(X1,k6_yellow_6(X0)) & l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) | ~v1_compts_1(X0)) & (! [X1] : (? [X2] : (r3_waybel_9(X0,X1,X2) & m1_subset_1(X2,u1_struct_0(X0))) | ~r2_hidden(X1,k6_yellow_6(X0)) | ~l1_waybel_0(X1,X0) | ~v7_waybel_0(X1) | ~v4_orders_2(X1) | v2_struct_0(X1)) | v1_compts_1(X0))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [nnf transformation 69]
134. ? [X0] : ((? [X1] : (! [X2] : (~r3_waybel_9(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) & r2_hidden(X1,k6_yellow_6(X0)) & l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) | ~v1_compts_1(X0)) & (! [X1] : (? [X2] : (r3_waybel_9(X0,X1,X2) & m1_subset_1(X2,u1_struct_0(X0))) | ~r2_hidden(X1,k6_yellow_6(X0)) | ~l1_waybel_0(X1,X0) | ~v7_waybel_0(X1) | ~v4_orders_2(X1) | v2_struct_0(X1)) | v1_compts_1(X0)) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [flattening 133]
135. ? [X0] : ((? [X1] : (! [X2] : (~r3_waybel_9(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) & r2_hidden(X1,k6_yellow_6(X0)) & l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) | ~v1_compts_1(X0)) & (! [X3] : (? [X4] : (r3_waybel_9(X0,X3,X4) & m1_subset_1(X4,u1_struct_0(X0))) | ~r2_hidden(X3,k6_yellow_6(X0)) | ~l1_waybel_0(X3,X0) | ~v7_waybel_0(X3) | ~v4_orders_2(X3) | v2_struct_0(X3)) | v1_compts_1(X0)) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [rectify 134]
136. ? [X0] : ((? [X1] : (! [X2] : (~r3_waybel_9(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) & r2_hidden(X1,k6_yellow_6(X0)) & l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) | ~v1_compts_1(X0)) & (! [X3] : (? [X4] : (r3_waybel_9(X0,X3,X4) & m1_subset_1(X4,u1_struct_0(X0))) | ~r2_hidden(X3,k6_yellow_6(X0)) | ~l1_waybel_0(X3,X0) | ~v7_waybel_0(X3) | ~v4_orders_2(X3) | v2_struct_0(X3)) | v1_compts_1(X0)) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ((? [X1] : (! [X2] : (~r3_waybel_9(sK0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(sK0))) & r2_hidden(X1,k6_yellow_6(sK0)) & l1_waybel_0(X1,sK0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) | ~v1_compts_1(sK0)) & (! [X3] : (? [X4] : (r3_waybel_9(sK0,X3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) | ~r2_hidden(X3,k6_yellow_6(sK0)) | ~l1_waybel_0(X3,sK0) | ~v7_waybel_0(X3) | ~v4_orders_2(X3) | v2_struct_0(X3)) | v1_compts_1(sK0)) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0)) [choice axiom]
137. ? [X1] : (! [X2] : (~r3_waybel_9(sK0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(sK0))) & r2_hidden(X1,k6_yellow_6(sK0)) & l1_waybel_0(X1,sK0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) => (! [X2] : (~r3_waybel_9(sK0,sK1,X2) | ~m1_subset_1(X2,u1_struct_0(sK0))) & r2_hidden(sK1,k6_yellow_6(sK0)) & l1_waybel_0(sK1,sK0) & v7_waybel_0(sK1) & v4_orders_2(sK1) & ~v2_struct_0(sK1)) [choice axiom]
138. ! [X3] : (? [X4] : (r3_waybel_9(sK0,X3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) => (r3_waybel_9(sK0,X3,sK2(X3)) & m1_subset_1(sK2(X3),u1_struct_0(sK0)))) [choice axiom]
139. ((! [X2] : (~r3_waybel_9(sK0,sK1,X2) | ~m1_subset_1(X2,u1_struct_0(sK0))) & r2_hidden(sK1,k6_yellow_6(sK0)) & l1_waybel_0(sK1,sK0) & v7_waybel_0(sK1) & v4_orders_2(sK1) & ~v2_struct_0(sK1)) | ~v1_compts_1(sK0)) & (! [X3] : ((r3_waybel_9(sK0,X3,sK2(X3)) & m1_subset_1(sK2(X3),u1_struct_0(sK0))) | ~r2_hidden(X3,k6_yellow_6(sK0)) | ~l1_waybel_0(X3,sK0) | ~v7_waybel_0(X3) | ~v4_orders_2(X3) | v2_struct_0(X3)) | v1_compts_1(sK0)) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0) [skolemisation 135,138,137,136]
142. ! [X1,X0] : (? [X2] : (r3_waybel_9(X0,X1,X2) & m1_subset_1(X2,u1_struct_0(X0))) => (r3_waybel_9(X0,X1,sK4(X0,X1)) & m1_subset_1(sK4(X0,X1),u1_struct_0(X0)))) [choice axiom]
143. ! [X0] : (! [X1] : ((r3_waybel_9(X0,X1,sK4(X0,X1)) & m1_subset_1(sK4(X0,X1),u1_struct_0(X0))) | ~l1_waybel_0(X1,X0) | ~v7_waybel_0(X1) | ~v4_orders_2(X1) | v2_struct_0(X1)) | ~v1_compts_1(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [skolemisation 108,142]
144. ! [X0] : (? [X1] : (! [X2] : (~r3_waybel_9(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) & r2_hidden(X1,k6_yellow_6(X0)) & l1_waybel_0(X1,X0) & v7_waybel_0(X1) & v4_orders_2(X1) & ~v2_struct_0(X1)) => (! [X2] : (~r3_waybel_9(X0,sK5(X0),X2) | ~m1_subset_1(X2,u1_struct_0(X0))) & r2_hidden(sK5(X0),k6_yellow_6(X0)) & l1_waybel_0(sK5(X0),X0) & v7_waybel_0(sK5(X0)) & v4_orders_2(sK5(X0)) & ~v2_struct_0(sK5(X0)))) [choice axiom]
145. ! [X0] : (v1_compts_1(X0) | (! [X2] : (~r3_waybel_9(X0,sK5(X0),X2) | ~m1_subset_1(X2,u1_struct_0(X0))) & r2_hidden(sK5(X0),k6_yellow_6(X0)) & l1_waybel_0(sK5(X0),X0) & v7_waybel_0(sK5(X0)) & v4_orders_2(sK5(X0)) & ~v2_struct_0(sK5(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [skolemisation 110,144]
162. ~v2_struct_0(sK0) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
163. v2_pre_topc(sK0) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
164. l1_pre_topc(sK0) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
165. m1_subset_1(sK2(X3),u1_struct_0(sK0)) | ~r2_hidden(X3,k6_yellow_6(sK0)) | ~l1_waybel_0(X3,sK0) | ~v7_waybel_0(X3) | ~v4_orders_2(X3) | v2_struct_0(X3) | v1_compts_1(sK0) [cnf transformation 139] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
166. r3_waybel_9(sK0,X3,sK2(X3)) | ~r2_hidden(X3,k6_yellow_6(sK0)) | ~l1_waybel_0(X3,sK0) | ~v7_waybel_0(X3) | ~v4_orders_2(X3) | v2_struct_0(X3) | v1_compts_1(sK0) [cnf transformation 139] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
167. ~v2_struct_0(sK1) | ~v1_compts_1(sK0) [cnf transformation 139] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
168. v4_orders_2(sK1) | ~v1_compts_1(sK0) [cnf transformation 139] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
169. v7_waybel_0(sK1) | ~v1_compts_1(sK0) [cnf transformation 139] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
170. l1_waybel_0(sK1,sK0) | ~v1_compts_1(sK0) [cnf transformation 139] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
172. ~r3_waybel_9(sK0,sK1,X2) | ~m1_subset_1(X2,u1_struct_0(sK0)) | ~v1_compts_1(sK0) [cnf transformation 139] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
203. m1_subset_1(sK4(X0,X1),u1_struct_0(X0)) | ~l1_waybel_0(X1,X0) | ~v7_waybel_0(X1) | ~v4_orders_2(X1) | v2_struct_0(X1) | ~v1_compts_1(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 143] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
204. r3_waybel_9(X0,X1,sK4(X0,X1)) | ~l1_waybel_0(X1,X0) | ~v7_waybel_0(X1) | ~v4_orders_2(X1) | v2_struct_0(X1) | ~v1_compts_1(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 143] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
205. ~v2_struct_0(sK5(X0)) | v1_compts_1(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 145] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
206. v4_orders_2(sK5(X0)) | v1_compts_1(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 145] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
207. v7_waybel_0(sK5(X0)) | v1_compts_1(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 145] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
208. l1_waybel_0(sK5(X0),X0) | v1_compts_1(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 145] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
209. r2_hidden(sK5(X0),k6_yellow_6(X0)) | v1_compts_1(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 145] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
210. ~r3_waybel_9(X0,sK5(X0),X2) | v1_compts_1(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 145] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
237. 1 <=> v1_compts_1(sK0) [avatar definition]
238. v1_compts_1(sK0) <- (1) [avatar component clause 237] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
239. ~v1_compts_1(sK0) <- (~1) [avatar component clause 237] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
241. 2 <=> ! [X2] : (~r3_waybel_9(sK0,sK1,X2) | ~m1_subset_1(X2,u1_struct_0(sK0))) [avatar definition]
242. ~r3_waybel_9(sK0,sK1,X2) | ~m1_subset_1(X2,u1_struct_0(sK0)) <- (2) [avatar component clause 241] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
243. ~1 | 2 [avatar split clause 172,241,237]
250. 4 <=> l1_waybel_0(sK1,sK0) [avatar definition]
252. l1_waybel_0(sK1,sK0) <- (4) [avatar component clause 250] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
253. ~1 | 4 [avatar split clause 170,250,237]
255. 5 <=> v7_waybel_0(sK1) [avatar definition]
257. v7_waybel_0(sK1) <- (5) [avatar component clause 255] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
258. ~1 | 5 [avatar split clause 169,255,237]
260. 6 <=> v4_orders_2(sK1) [avatar definition]
262. v4_orders_2(sK1) <- (6) [avatar component clause 260] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
263. ~1 | 6 [avatar split clause 168,260,237]
265. 7 <=> v2_struct_0(sK1) [avatar definition]
267. ~v2_struct_0(sK1) <- (~7) [avatar component clause 265] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
268. ~1 | ~7 [avatar split clause 167,265,237]
270. 8 <=> ! [X3] : (r3_waybel_9(sK0,X3,sK2(X3)) | v2_struct_0(X3) | ~v4_orders_2(X3) | ~v7_waybel_0(X3) | ~l1_waybel_0(X3,sK0) | ~r2_hidden(X3,k6_yellow_6(sK0))) [avatar definition]
271. r3_waybel_9(sK0,X3,sK2(X3)) | v2_struct_0(X3) | ~v4_orders_2(X3) | ~v7_waybel_0(X3) | ~l1_waybel_0(X3,sK0) | ~r2_hidden(X3,k6_yellow_6(sK0)) <- (8) [avatar component clause 270] {a:0,w:18,wCS:1800,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
272. 1 | 8 [avatar split clause 166,270,237]
274. 9 <=> ! [X3] : (m1_subset_1(sK2(X3),u1_struct_0(sK0)) | v2_struct_0(X3) | ~v4_orders_2(X3) | ~v7_waybel_0(X3) | ~l1_waybel_0(X3,sK0) | ~r2_hidden(X3,k6_yellow_6(sK0))) [avatar definition]
275. m1_subset_1(sK2(X3),u1_struct_0(sK0)) | v2_struct_0(X3) | ~v4_orders_2(X3) | ~v7_waybel_0(X3) | ~l1_waybel_0(X3,sK0) | ~r2_hidden(X3,k6_yellow_6(sK0)) <- (9) [avatar component clause 274] {a:0,w:18,wCS:1800,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
276. 1 | 9 [avatar split clause 165,274,237]
278. 10 <=> l1_pre_topc(sK0) [avatar definition]
280. l1_pre_topc(sK0) <- (10) [avatar component clause 278] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
281. 10 [avatar split clause 164,278]
283. 11 <=> v2_pre_topc(sK0) [avatar definition]
285. v2_pre_topc(sK0) <- (11) [avatar component clause 283] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
286. 11 [avatar split clause 163,283]
288. 12 <=> v2_struct_0(sK0) [avatar definition]
290. ~v2_struct_0(sK0) <- (~12) [avatar component clause 288] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
291. ~12 [avatar split clause 162,288]
515. v1_compts_1(sK0) | ~m1_subset_1(sK2(sK5(sK0)),u1_struct_0(sK0)) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) | v2_struct_0(sK5(sK0)) | ~v4_orders_2(sK5(sK0)) | ~v7_waybel_0(sK5(sK0)) | ~l1_waybel_0(sK5(sK0),sK0) | ~r2_hidden(sK5(sK0),k6_yellow_6(sK0)) <- (8) [resolution 210,271] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
522. 47 <=> r2_hidden(sK5(sK0),k6_yellow_6(sK0)) [avatar definition]
524. ~r2_hidden(sK5(sK0),k6_yellow_6(sK0)) <- (~47) [avatar component clause 522] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
526. 48 <=> l1_waybel_0(sK5(sK0),sK0) [avatar definition]
528. ~l1_waybel_0(sK5(sK0),sK0) <- (~48) [avatar component clause 526] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
530. 49 <=> v7_waybel_0(sK5(sK0)) [avatar definition]
532. ~v7_waybel_0(sK5(sK0)) <- (~49) [avatar component clause 530] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
534. 50 <=> v4_orders_2(sK5(sK0)) [avatar definition]
536. ~v4_orders_2(sK5(sK0)) <- (~50) [avatar component clause 534] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
538. 51 <=> v2_struct_0(sK5(sK0)) [avatar definition]
540. v2_struct_0(sK5(sK0)) <- (51) [avatar component clause 538] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
543. v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~47) [resolution 524,209] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
545. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, ~47) [subsumption resolution 543,239] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
546. ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, 10, ~47) [subsumption resolution 545,280] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
547. v2_struct_0(sK0) <- (~1, 10, 11, ~47) [subsumption resolution 546,285] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
548. $false <- (~1, 10, 11, ~12, ~47) [subsumption resolution 547,290] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
549. 1 | ~10 | ~11 | 12 | 47 [avatar contradiction clause 548]
555. v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~48) [resolution 528,208] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
557. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, ~48) [subsumption resolution 555,239] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
558. ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, 10, ~48) [subsumption resolution 557,280] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
559. v2_struct_0(sK0) <- (~1, 10, 11, ~48) [subsumption resolution 558,285] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
560. $false <- (~1, 10, 11, ~12, ~48) [subsumption resolution 559,290] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
561. 1 | ~10 | ~11 | 12 | 48 [avatar contradiction clause 560]
563. v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~49) [resolution 532,207] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
565. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, ~49) [subsumption resolution 563,239] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
566. ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, 10, ~49) [subsumption resolution 565,280] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
567. v2_struct_0(sK0) <- (~1, 10, 11, ~49) [subsumption resolution 566,285] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
568. $false <- (~1, 10, 11, ~12, ~49) [subsumption resolution 567,290] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
569. 1 | ~10 | ~11 | 12 | 49 [avatar contradiction clause 568]
571. v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~50) [resolution 536,206] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
573. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, ~50) [subsumption resolution 571,239] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
574. ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, 10, ~50) [subsumption resolution 573,280] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
575. v2_struct_0(sK0) <- (~1, 10, 11, ~50) [subsumption resolution 574,285] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
576. $false <- (~1, 10, 11, ~12, ~50) [subsumption resolution 575,290] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
577. 1 | ~10 | ~11 | 12 | 50 [avatar contradiction clause 576]
579. v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (51) [resolution 540,205] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
581. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, 51) [subsumption resolution 579,239] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
582. ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~1, 10, 51) [subsumption resolution 581,280] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
583. v2_struct_0(sK0) <- (~1, 10, 11, 51) [subsumption resolution 582,285] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
584. $false <- (~1, 10, 11, ~12, 51) [subsumption resolution 583,290] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
585. 1 | ~10 | ~11 | 12 | ~51 [avatar contradiction clause 584]
586. v1_compts_1(sK0) | ~m1_subset_1(sK2(sK5(sK0)),u1_struct_0(sK0)) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) | v2_struct_0(sK5(sK0)) | ~v4_orders_2(sK5(sK0)) | ~v7_waybel_0(sK5(sK0)) | ~l1_waybel_0(sK5(sK0),sK0) | ~r2_hidden(sK5(sK0),k6_yellow_6(sK0)) <- (8, 10) [subsumption resolution 515,280] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
587. v1_compts_1(sK0) | ~m1_subset_1(sK2(sK5(sK0)),u1_struct_0(sK0)) | v2_struct_0(sK0) | v2_struct_0(sK5(sK0)) | ~v4_orders_2(sK5(sK0)) | ~v7_waybel_0(sK5(sK0)) | ~l1_waybel_0(sK5(sK0),sK0) | ~r2_hidden(sK5(sK0),k6_yellow_6(sK0)) <- (8, 10, 11) [subsumption resolution 586,285] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
588. v1_compts_1(sK0) | ~m1_subset_1(sK2(sK5(sK0)),u1_struct_0(sK0)) | v2_struct_0(sK5(sK0)) | ~v4_orders_2(sK5(sK0)) | ~v7_waybel_0(sK5(sK0)) | ~l1_waybel_0(sK5(sK0),sK0) | ~r2_hidden(sK5(sK0),k6_yellow_6(sK0)) <- (8, 10, 11, ~12) [subsumption resolution 587,290] {a:1,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
589. v1_compts_1(sK0) | v2_struct_0(sK5(sK0)) | ~v4_orders_2(sK5(sK0)) | ~v7_waybel_0(sK5(sK0)) | ~l1_waybel_0(sK5(sK0),sK0) | ~r2_hidden(sK5(sK0),k6_yellow_6(sK0)) <- (8, 9, 10, 11, ~12) [subsumption resolution 588,275] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
590. ~47 | ~48 | ~49 | ~50 | 51 | 1 | ~8 | ~9 | ~10 | ~11 | 12 [avatar split clause 589,288,283,278,274,270,237,538,534,530,526,522]
623. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) | ~l1_waybel_0(sK1,sK0) | ~v7_waybel_0(sK1) | ~v4_orders_2(sK1) | v2_struct_0(sK1) | ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2) [resolution 242,204] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
634. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) | ~v7_waybel_0(sK1) | ~v4_orders_2(sK1) | v2_struct_0(sK1) | ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2, 4) [subsumption resolution 623,252] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
635. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) | ~v4_orders_2(sK1) | v2_struct_0(sK1) | ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2, 4, 5) [subsumption resolution 634,257] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
636. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) | v2_struct_0(sK1) | ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2, 4, 5, 6) [subsumption resolution 635,262] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
637. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) | ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2, 4, 5, 6, ~7) [subsumption resolution 636,267] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
638. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (1, 2, 4, 5, 6, ~7) [subsumption resolution 637,238] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
639. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (1, 2, 4, 5, 6, ~7, 10) [subsumption resolution 638,280] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
640. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) | v2_struct_0(sK0) <- (1, 2, 4, 5, 6, ~7, 10, 11) [subsumption resolution 639,285] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
641. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) <- (1, 2, 4, 5, 6, ~7, 10, 11, ~12) [subsumption resolution 640,290] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
643. 58 <=> m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) [avatar definition]
645. ~m1_subset_1(sK4(sK0,sK1),u1_struct_0(sK0)) <- (~58) [avatar component clause 643] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
646. ~58 | ~1 | ~2 | ~4 | ~5 | ~6 | 7 | ~10 | ~11 | 12 [avatar split clause 641,288,283,278,265,260,255,250,241,237,643]
671. ~l1_waybel_0(sK1,sK0) | ~v7_waybel_0(sK1) | ~v4_orders_2(sK1) | v2_struct_0(sK1) | ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (~58) [resolution 645,203] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
673. ~v7_waybel_0(sK1) | ~v4_orders_2(sK1) | v2_struct_0(sK1) | ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, ~58) [subsumption resolution 671,252] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
674. ~v4_orders_2(sK1) | v2_struct_0(sK1) | ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 5, ~58) [subsumption resolution 673,257] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
675. v2_struct_0(sK1) | ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 5, 6, ~58) [subsumption resolution 674,262] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
676. ~v1_compts_1(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 5, 6, ~7, ~58) [subsumption resolution 675,267] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
677. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (1, 4, 5, 6, ~7, ~58) [subsumption resolution 676,238] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
678. ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (1, 4, 5, 6, ~7, 10, ~58) [subsumption resolution 677,280] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
679. v2_struct_0(sK0) <- (1, 4, 5, 6, ~7, 10, 11, ~58) [subsumption resolution 678,285] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
680. $false <- (1, 4, 5, 6, ~7, 10, 11, ~12, ~58) [subsumption resolution 679,290] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
681. ~1 | ~4 | ~5 | ~6 | 7 | ~10 | ~11 | 12 | 58 [avatar contradiction clause 680]
682. $false [avatar sat refutation 243,253,258,263,268,272,276,281,286,291,549,561,569,577,585,590,646,681] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:568,thDist:-568}
% SZS output end Proof for yellow19__t34_yellow19
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 57
% >>> Preprocessing
% Introduced skolems: 14
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 74
% Generated clauses: 276
% Active clauses: 97
% Passive clauses: 120
% Final active clauses: 88
% Final passive clauses: 15
% Main loop iterations started: 105
% 
% >>> Simplifying Inferences
% Duplicate literals: 10
% Fw subsumption resolutions: 57
% Fw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 17
% 
% >>> Generating Inferences
% Binary resolution: 65
% Unit resulting resolution: 7
% Forward superposition: 2
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 20
% Split components: 47
% Sat splitting refutations: 14
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 67
% SAT solver unit clauses: 17
% SAT solver binary clauses: 19
% 
% Memory used [KB]: 5373
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% superposition: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

