i: [1,69,0,2,1,1,0,0] t41_relat_1
a: [3,115,0,2,1,69]
s: 115
i: [1,71,0,3,1,1,0,0] t41_relat_1
a: [3,101,0,3,1,71]
s: 101
i: [1,70,0,8,2,1,0,0] t41_relat_1
a: [3,110,0,4,1,70]
s: 110
i: [1,72,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,120,0,2,1,72]
s: 120
i: [1,73,0,2,1,0,0,1] fc1_xboole_0
a: [3,125,0,2,1,73]
s: 125
i: [1,91,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,135,0,2,1,91]
s: 135
i: [1,92,0,2,1,0,0,2] rc2_xboole_0
a: [3,140,0,2,1,92]
s: 140
i: [1,93,0,2,1,0,0,2] rc1_xboole_0
a: [3,145,0,2,1,93]
s: 145
i: [1,95,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,150,0,2,1,95]
s: 150
i: [1,94,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,155,0,2,1,94]
s: 155
i: [1,97,0,2,1,0,0,1] rc1_relat_1
a: [3,160,0,2,1,97]
s: 160
i: [1,96,0,2,1,0,0,1] rc1_relat_1
a: [3,165,0,2,1,96]
s: 165
i: [1,74,0,3,1,0,0,255] d2_xboole_0
a: [3,130,0,3,1,74]
s: 130
r: 120
i: [1,75,0,4,2,0,0,255] cc2_zfmisc_1
s: 75
i: [1,76,0,4,2,0,0,255] cc1_zfmisc_1
s: 76
i: [1,77,0,4,2,0,0,1] cc1_relat_1
s: 77
d: [2,168,1,2,1,1,40,75,145]
a: [3,173,1,2,1,168]
s: 173
i: [1,89,0,4,1,0,0,255] existence_m1_subset_1
s: 89
d: [2,169,1,2,1,1,40,75,125]
a: [3,178,1,2,1,169]
s: 178
i: [1,78,0,5,2,0,0,1] t6_boole
s: 78
i: [1,79,0,5,2,0,0,1] fc11_relat_1
s: 79
d: [2,180,1,2,1,1,40,77,145]
a: [3,185,1,2,1,180]
s: 185
i: [1,80,0,5,2,0,0,1] fc10_relat_1
s: 80
d: [2,181,1,2,1,1,40,77,125]
a: [3,190,1,2,1,181]
s: 190
i: [1,81,0,5,2,0,0,1] fc5_xtuple_0
s: 81
i: [1,82,0,5,2,0,0,1] fc4_xtuple_0
s: 82
i: [1,88,0,5,2,0,0,255] t7_boole
s: 88
i: [1,85,0,6,2,0,0,255] t1_subset
s: 85
d: [2,195,1,5,2,0,40,79,77]
s: 195
i: [1,90,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 90
i: [1,83,0,7,3,0,0,1] fc9_relat_1
s: 83
d: [2,203,1,6,3,1,44,83,110]
d: [2,204,1,4,2,2,21,203,125]
d: [2,205,1,2,1,3,21,204,115]
a: [3,209,1,2,1,205]
s: 209
d: [2,211,2,0,0,2,52,101,209,78]
e: 211
r: 110
r: 209
a: [3,106,0,4,1,70]
s: 106
a: [3,208,1,2,1,205]
s: 208
a: [3,109,0,4,1,70]
s: 109
i: [1,84,0,7,3,0,0,1] fc8_relat_1
s: 84
d: [2,217,1,6,3,1,44,84,106]
d: [2,218,1,4,2,2,21,217,125]
d: [2,219,1,2,1,3,21,218,115]
d: [2,220,1,0,0,4,21,219,208]
e: 220
f: 220,211,205,73,69,70,71
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t41_relat_1
% SZS output start Proof for relat_1__t41_relat_1
1. ! [X0] : (v1_relat_1(X0) => ((k1_xboole_0 = k10_xtuple_0(X0) | k9_xtuple_0(X0) = k1_xboole_0) => k1_xboole_0 = X0)) [input t41_relat_1]
2. ~! [X0] : (v1_relat_1(X0) => ((k1_xboole_0 = k10_xtuple_0(X0) | k9_xtuple_0(X0) = k1_xboole_0) => k1_xboole_0 = X0)) [negated conjecture 1]
18. ! [X0] : ((v1_relat_1(X0) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k9_xtuple_0(X0))) [input fc8_relat_1]
19. ! [X0] : ((v1_relat_1(X0) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k10_xtuple_0(X0))) [input fc9_relat_1]
28. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
29. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
35. ? [X0] : ((k1_xboole_0 != X0 & (k1_xboole_0 = k10_xtuple_0(X0) | k9_xtuple_0(X0) = k1_xboole_0)) & v1_relat_1(X0)) [ennf transformation 2]
36. ? [X0] : (k1_xboole_0 != X0 & (k1_xboole_0 = k10_xtuple_0(X0) | k9_xtuple_0(X0) = k1_xboole_0) & v1_relat_1(X0)) [flattening 35]
40. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 29]
45. ! [X0] : (~v1_xboole_0(k10_xtuple_0(X0)) | (~v1_relat_1(X0) | v1_xboole_0(X0))) [ennf transformation 19]
46. ! [X0] : (~v1_xboole_0(k10_xtuple_0(X0)) | ~v1_relat_1(X0) | v1_xboole_0(X0)) [flattening 45]
47. ! [X0] : (~v1_xboole_0(k9_xtuple_0(X0)) | (~v1_relat_1(X0) | v1_xboole_0(X0))) [ennf transformation 18]
48. ! [X0] : (~v1_xboole_0(k9_xtuple_0(X0)) | ~v1_relat_1(X0) | v1_xboole_0(X0)) [flattening 47]
55. ? [X0] : (k1_xboole_0 != X0 & (k1_xboole_0 = k10_xtuple_0(X0) | k9_xtuple_0(X0) = k1_xboole_0) & v1_relat_1(X0)) => (k1_xboole_0 != sK0 & (k1_xboole_0 = k10_xtuple_0(sK0) | k1_xboole_0 = k9_xtuple_0(sK0)) & v1_relat_1(sK0)) [choice axiom]
56. k1_xboole_0 != sK0 & (k1_xboole_0 = k10_xtuple_0(sK0) | k1_xboole_0 = k9_xtuple_0(sK0)) & v1_relat_1(sK0) [skolemisation 36,55]
69. v1_relat_1(sK0) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. k1_xboole_0 = k10_xtuple_0(sK0) | k1_xboole_0 = k9_xtuple_0(sK0) [cnf transformation 56] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
71. k1_xboole_0 != sK0 [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
73. v1_xboole_0(k1_xboole_0) [cnf transformation 28] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
78. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 40] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
83. ~v1_xboole_0(k10_xtuple_0(X0)) | ~v1_relat_1(X0) | v1_xboole_0(X0) [cnf transformation 46] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
84. ~v1_xboole_0(k9_xtuple_0(X0)) | ~v1_relat_1(X0) | v1_xboole_0(X0) [cnf transformation 48] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
99. 1 <=> k1_xboole_0 = sK0 [avatar definition]
101. k1_xboole_0 != sK0 <- (~1) [avatar component clause 99] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
102. ~1 [avatar split clause 71,99]
104. 2 <=> k1_xboole_0 = k9_xtuple_0(sK0) [avatar definition]
106. k1_xboole_0 = k9_xtuple_0(sK0) <- (2) [avatar component clause 104] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
108. 3 <=> k1_xboole_0 = k10_xtuple_0(sK0) [avatar definition]
110. k1_xboole_0 = k10_xtuple_0(sK0) <- (3) [avatar component clause 108] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
111. 2 | 3 [avatar split clause 70,108,104]
113. 4 <=> v1_relat_1(sK0) [avatar definition]
115. v1_relat_1(sK0) <- (4) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. 4 [avatar split clause 69,113]
123. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
125. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 123] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
126. 6 [avatar split clause 73,123]
203. ~v1_xboole_0(k1_xboole_0) | ~v1_relat_1(sK0) | v1_xboole_0(sK0) <- (3) [superposition 83,110] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
204. ~v1_relat_1(sK0) | v1_xboole_0(sK0) <- (3, 6) [subsumption resolution 203,125] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
205. v1_xboole_0(sK0) <- (3, 4, 6) [subsumption resolution 204,115] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
207. 19 <=> v1_xboole_0(sK0) [avatar definition]
208. ~v1_xboole_0(sK0) <- (~19) [avatar component clause 207] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
209. v1_xboole_0(sK0) <- (19) [avatar component clause 207] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
210. 19 | ~3 | ~4 | ~6 [avatar split clause 205,123,113,108,207]
211. $false <- (~1, 19) [unit resulting resolution 101,209,78] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
214. 1 | ~19 [avatar contradiction clause 211]
217. ~v1_xboole_0(k1_xboole_0) | ~v1_relat_1(sK0) | v1_xboole_0(sK0) <- (2) [superposition 84,106] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
218. ~v1_relat_1(sK0) | v1_xboole_0(sK0) <- (2, 6) [subsumption resolution 217,125] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
219. v1_xboole_0(sK0) <- (2, 4, 6) [subsumption resolution 218,115] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
220. $false <- (2, 4, 6, ~19) [subsumption resolution 219,208] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
221. ~2 | ~4 | ~6 | 19 [avatar contradiction clause 220]
222. $false [avatar sat refutation 102,111,116,126,210,214,221] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:71,thDist:-71}
% SZS output end Proof for relat_1__t41_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 78
% Active clauses: 33
% Passive clauses: 42
% Final active clauses: 30
% Final passive clauses: 6
% Main loop iterations started: 36
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 12
% Unit resulting resolution: 1
% Forward superposition: 5
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 22
% SAT solver unit clauses: 12
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

