i: [1,90,0,2,1,1,0,0] t36_mssubfam
a: [3,175,0,2,1,90]
s: 175
i: [1,88,0,2,1,1,0,0] t36_mssubfam
a: [3,185,0,2,1,88]
s: 185
i: [1,86,0,2,1,1,0,0] t36_mssubfam
a: [3,195,0,2,1,86]
s: 195
i: [1,84,0,2,1,1,0,0] t36_mssubfam
a: [3,205,0,2,1,84]
s: 205
i: [1,91,0,3,1,1,0,0] t36_mssubfam
a: [3,170,0,3,1,91]
s: 170
i: [1,89,0,3,1,1,0,0] t36_mssubfam
a: [3,180,0,3,1,89]
s: 180
i: [1,87,0,3,1,1,0,0] t36_mssubfam
a: [3,190,0,3,1,87]
s: 190
i: [1,85,0,3,1,1,0,0] t36_mssubfam
a: [3,200,0,3,1,85]
s: 200
i: [1,92,0,6,1,1,0,0] t36_mssubfam
a: [3,165,0,6,1,92]
s: 165
i: [1,151,0,2,1,0,0,2] rc2_relat_1
a: [3,210,0,2,1,151]
s: 210
i: [1,150,0,2,1,0,0,2] rc2_relat_1
a: [3,215,0,2,1,150]
s: 215
i: [1,154,0,2,1,0,0,2] rc3_funct_1
a: [3,220,0,2,1,154]
s: 220
i: [1,153,0,2,1,0,0,2] rc3_funct_1
a: [3,225,0,2,1,153]
s: 225
i: [1,152,0,2,1,0,0,2] rc3_funct_1
a: [3,230,0,2,1,152]
s: 230
i: [1,156,0,2,1,0,0,1] rc1_funct_1
a: [3,235,0,2,1,156]
s: 235
i: [1,155,0,2,1,0,0,1] rc1_funct_1
a: [3,240,0,2,1,155]
s: 240
i: [1,93,0,9,1,1,0,0] t36_mssubfam
a: [3,160,0,9,1,93]
s: 160
i: [1,99,0,3,1,0,0,2] rc9_funct_1
s: 99
i: [1,100,0,3,1,0,0,2] rc9_funct_1
s: 100
i: [1,102,0,3,1,0,0,2] rc9_funct_1
s: 102
i: [1,103,0,3,1,0,0,2] rc4_funcop_1
s: 103
i: [1,104,0,3,1,0,0,2] rc4_funcop_1
s: 104
i: [1,106,0,3,1,0,0,2] rc4_funcop_1
s: 106
i: [1,108,0,3,1,0,0,2] rc3_pboole
s: 108
i: [1,109,0,3,1,0,0,2] rc3_pboole
s: 109
i: [1,111,0,3,1,0,0,2] rc3_pboole
s: 111
i: [1,113,0,3,1,0,0,1] rc1_pboole
s: 113
i: [1,115,0,3,1,0,0,1] rc1_pboole
s: 115
i: [1,101,0,4,1,0,0,2] rc9_funct_1
s: 101
i: [1,105,0,4,1,0,0,2] rc4_funcop_1
s: 105
i: [1,107,0,4,1,0,0,2] rc4_funcop_1
s: 107
i: [1,110,0,4,1,0,0,2] rc3_pboole
s: 110
i: [1,112,0,4,1,0,0,2] rc3_pboole
s: 112
i: [1,114,0,4,1,0,0,1] rc1_pboole
s: 114
i: [1,116,0,4,1,0,0,1] rc1_pboole
s: 116
i: [1,125,0,14,5,0,0,1] fc1_mboolean
s: 125
i: [1,126,0,14,5,0,0,1] fc1_mboolean
s: 126
i: [1,128,0,14,5,0,0,1] fc1_mboolean
s: 128
i: [1,130,0,14,5,0,0,1] dt_k1_mboolean
s: 130
i: [1,132,0,14,5,0,0,1] dt_k1_mboolean
s: 132
i: [1,140,0,14,5,0,0,1] rc2_mssubfam
s: 140
i: [1,141,0,14,5,0,0,1] rc2_mssubfam
s: 141
i: [1,143,0,14,5,0,0,1] rc2_mssubfam
s: 143
i: [1,127,0,15,5,0,0,1] fc1_mboolean
s: 127
i: [1,129,0,15,5,0,0,1] fc1_mboolean
s: 129
i: [1,131,0,15,5,0,0,1] dt_k1_mboolean
s: 131
i: [1,133,0,15,5,0,0,1] dt_k1_mboolean
s: 133
i: [1,142,0,15,5,0,0,1] rc2_mssubfam
s: 142
i: [1,144,0,15,5,0,0,1] rc2_mssubfam
s: 144
i: [1,118,0,16,6,0,0,1] rc5_pboole
s: 118
i: [1,119,0,16,6,0,0,1] rc5_pboole
s: 119
i: [1,121,0,16,6,0,0,1] rc5_pboole
s: 121
i: [1,134,0,16,6,0,0,1] dt_m3_pboole
s: 134
d: [2,242,1,20,5,1,40,134,165]
a: [3,246,1,4,1,242]
s: 246
d: [2,264,2,0,0,5,52,185,175,180,170,246,125]
e: 264
d: [2,265,2,10,4,1,40,246,125]
d: [2,267,2,7,3,2,21,265,170]
d: [2,268,2,5,2,3,21,267,175]
d: [2,269,2,2,1,4,21,268,180]
d: [2,270,2,0,0,5,21,269,185]
e: 270
r: 246
i: [1,136,0,16,6,0,0,1] dt_m3_pboole
s: 136
i: [1,138,0,16,5,0,0,1] existence_m3_pboole
s: 138
a: [3,245,1,4,1,242]
s: 245
i: [1,120,0,17,6,0,0,1] rc5_pboole
s: 120
a: [3,250,1,5,1,242]
s: 250
d: [2,282,2,0,0,5,52,185,175,180,170,250,127]
e: 282
d: [2,283,2,10,4,1,40,250,127]
d: [2,285,2,7,3,2,21,283,170]
d: [2,286,2,5,2,3,21,285,175]
d: [2,287,2,2,1,4,21,286,180]
d: [2,288,2,0,0,5,21,287,185]
e: 288
r: 250
i: [1,122,0,17,6,0,0,1] rc5_pboole
s: 122
i: [1,123,0,17,5,0,0,255] redefinition_k1_mssubfam
s: 123
a: [3,254,1,4,1,242]
s: 254
d: [2,290,2,0,0,5,52,185,175,180,170,254,128]
e: 290
d: [2,291,2,10,4,1,40,254,128]
d: [2,293,2,7,3,2,21,291,170]
d: [2,294,2,5,2,3,21,293,175]
d: [2,295,2,2,1,4,21,294,180]
d: [2,296,2,0,0,5,21,295,185]
e: 296
r: 254
i: [1,135,0,17,6,0,0,1] dt_m3_pboole
s: 135
a: [3,253,1,4,1,242]
s: 253
i: [1,137,0,17,6,0,0,1] dt_m3_pboole
s: 137
i: [1,117,0,18,6,0,0,1] rc5_pboole
s: 117
a: [3,249,1,5,1,242]
s: 249
i: [1,124,0,18,5,0,0,255] dt_k1_mssubfam
s: 124
a: [3,258,1,5,1,242]
s: 258
d: [2,347,2,0,0,5,52,185,175,180,170,258,129]
e: 347
d: [2,348,2,10,4,1,40,258,129]
d: [2,350,2,7,3,2,21,348,170]
d: [2,351,2,5,2,3,21,350,175]
d: [2,352,2,2,1,4,21,351,180]
d: [2,353,2,0,0,5,21,352,185]
e: 353
r: 258
i: [1,139,0,18,5,0,0,1] rc2_mssubfam
s: 139
i: [1,145,0,24,9,0,0,2] reflexivity_r2_pboole
s: 145
a: [3,262,1,2,1,242]
s: 262
i: [1,146,0,25,9,0,0,1] dt_k4_pboole
s: 146
d: [2,272,1,20,5,1,40,136,165]
a: [3,276,1,2,1,272]
s: 276
i: [1,148,0,25,9,0,0,1] dt_k4_pboole
s: 148
i: [1,147,0,26,9,0,0,1] dt_k4_pboole
s: 147
d: [2,298,1,21,5,1,40,135,165]
d: [2,301,1,17,4,2,21,298,245]
a: [3,305,1,3,1,301]
s: 305
i: [1,149,0,26,9,0,0,1] dt_k4_pboole
s: 149
d: [2,307,1,21,5,1,40,137,165]
d: [2,310,1,17,4,2,21,307,253]
d: [2,311,1,13,3,3,21,310,245]
a: [3,315,1,3,1,311]
s: 315
i: [1,94,0,27,9,0,0,1] t56_pboole
s: 94
i: [1,95,0,28,10,0,0,2] d18_pboole
s: 95
a: [3,257,1,5,1,242]
s: 257
d: [2,359,0,26,9,0,21,95,134]
d: [2,360,0,24,8,0,21,359,136]
d: [2,361,0,21,7,0,21,360,135]
d: [2,362,0,18,6,0,21,361,137]
s: 362
d: [2,364,1,24,5,1,40,362,165]
d: [2,374,1,19,4,2,21,364,257]
d: [2,375,1,15,3,3,21,374,253]
d: [2,376,1,10,2,4,21,375,249]
d: [2,377,1,6,1,5,21,376,245]
a: [3,381,1,6,1,377]
s: 381
i: [1,96,0,28,10,0,0,2] d18_pboole
s: 96
i: [1,98,0,42,15,0,0,2] t13_pboole
s: 98
d: [2,398,2,48,14,1,40,98,381]
d: [2,405,2,43,13,2,21,398,257]
d: [2,406,2,39,12,3,21,405,253]
d: [2,407,2,34,11,4,21,406,249]
d: [2,408,2,30,10,5,21,407,245]
d: [2,409,2,27,9,6,21,408,315]
d: [2,410,2,25,8,7,21,409,276]
d: [2,411,2,22,7,8,21,410,305]
d: [2,412,2,20,6,9,21,411,262]
s: 412
i: [1,97,0,44,14,0,0,1] t53_pboole
s: 97
d: [2,414,3,48,14,9,40,412,96]
d: [2,415,3,38,10,9,18,414]
d: [2,421,3,33,9,9,21,415,257]
d: [2,422,3,29,8,9,21,421,253]
d: [2,423,3,24,7,9,21,422,249]
d: [2,424,3,20,6,9,21,423,245]
s: 424
d: [2,279,1,24,9,0,40,138,134]
d: [2,280,1,14,5,0,18,279]
s: 280
d: [2,278,1,24,9,0,40,138,136]
d: [2,281,1,14,5,0,18,278]
s: 281
d: [2,443,4,29,5,10,40,424,160]
a: [3,452,4,7,1,443]
s: 452
d: [2,470,5,0,0,9,52,262,205,276,195,305,315,200,190,452,94]
e: 470
d: [2,471,5,20,8,1,40,452,94]
d: [2,473,5,17,7,2,21,471,190]
d: [2,474,5,15,6,3,21,473,195]
d: [2,475,5,12,5,4,21,474,200]
d: [2,476,5,10,4,5,21,475,205]
d: [2,477,5,7,3,6,21,476,315]
d: [2,478,5,5,2,7,21,477,276]
d: [2,479,5,2,1,8,21,478,305]
d: [2,480,5,0,0,9,21,479,262]
e: 480
r: 452
d: [2,327,1,32,9,0,40,124,136]
d: [2,339,1,27,8,0,21,327,129]
d: [2,340,1,23,7,0,21,339,128]
d: [2,341,1,18,6,0,21,340,127]
d: [2,342,1,14,5,0,21,341,125]
s: 342
a: [3,456,4,5,1,443]
s: 456
d: [2,484,5,0,0,9,52,262,205,276,195,305,315,200,190,456,146]
e: 484
d: [2,485,5,20,8,1,40,456,146]
d: [2,487,5,17,7,2,21,485,190]
d: [2,488,5,15,6,3,21,487,195]
d: [2,489,5,12,5,4,21,488,200]
d: [2,490,5,10,4,5,21,489,205]
d: [2,491,5,7,3,6,21,490,315]
d: [2,492,5,5,2,7,21,491,276]
d: [2,493,5,2,1,8,21,492,305]
d: [2,494,5,0,0,9,21,493,262]
e: 494
r: 456
d: [2,328,1,32,9,0,40,124,134]
d: [2,343,1,27,8,0,21,328,129]
d: [2,344,1,23,7,0,21,343,128]
d: [2,345,1,18,6,0,21,344,127]
d: [2,346,1,14,5,0,21,345,125]
s: 346
d: [2,299,1,25,9,0,40,135,138]
d: [2,300,1,15,5,0,18,299]
s: 300
a: [3,455,4,5,1,443]
s: 455
d: [2,308,1,25,9,0,40,137,138]
d: [2,309,1,15,5,0,18,308]
s: 309
a: [3,460,4,6,1,443]
s: 460
d: [2,498,5,0,0,9,52,262,205,276,195,305,315,200,190,460,147]
e: 498
d: [2,499,5,20,8,1,40,460,147]
d: [2,501,5,17,7,2,21,499,190]
d: [2,502,5,15,6,3,21,501,195]
d: [2,503,5,12,5,4,21,502,200]
d: [2,504,5,10,4,5,21,503,205]
d: [2,505,5,7,3,6,21,504,315]
d: [2,506,5,5,2,7,21,505,276]
d: [2,507,5,2,1,8,21,506,305]
d: [2,508,5,0,0,9,21,507,262]
e: 508
r: 460
d: [2,325,1,33,9,0,40,124,137]
d: [2,331,1,28,8,0,21,325,129]
d: [2,332,1,24,7,0,21,331,128]
d: [2,333,1,19,6,0,21,332,127]
d: [2,334,1,15,5,0,21,333,125]
s: 334
d: [2,326,1,33,9,0,40,124,135]
d: [2,335,1,28,8,0,21,326,129]
d: [2,336,1,24,7,0,21,335,128]
d: [2,337,1,19,6,0,21,336,127]
d: [2,338,1,15,5,0,21,337,125]
s: 338
a: [3,464,4,5,1,443]
s: 464
d: [2,514,5,0,0,9,52,262,205,276,195,305,315,200,190,464,148]
e: 514
d: [2,515,5,20,8,1,40,464,148]
d: [2,517,5,17,7,2,21,515,190]
d: [2,518,5,15,6,3,21,517,195]
d: [2,519,5,12,5,4,21,518,200]
d: [2,520,5,10,4,5,21,519,205]
d: [2,521,5,7,3,6,21,520,315]
d: [2,522,5,5,2,7,21,521,276]
d: [2,523,5,2,1,8,21,522,305]
d: [2,524,5,0,0,9,21,523,262]
e: 524
r: 464
d: [2,366,1,26,9,0,40,362,138]
d: [2,368,1,16,5,0,18,366]
s: 368
a: [3,463,4,5,1,443]
s: 463
d: [2,329,1,28,9,0,44,124,123]
d: [2,330,1,18,5,0,18,329]
s: 330
d: [2,365,1,28,10,0,40,362,117]
d: [2,369,1,18,6,0,18,365]
s: 369
a: [3,459,4,6,1,443]
s: 459
d: [2,363,1,36,9,0,40,362,124]
d: [2,370,1,31,8,0,21,363,129]
d: [2,371,1,27,7,0,21,370,128]
d: [2,372,1,22,6,0,21,371,127]
d: [2,373,1,18,5,0,21,372,125]
s: 373
a: [3,468,4,6,1,443]
s: 468
d: [2,559,5,0,0,9,52,262,205,276,195,305,315,200,190,468,149]
e: 559
d: [2,560,5,20,8,1,40,468,149]
d: [2,562,5,17,7,2,21,560,190]
d: [2,563,5,15,6,3,21,562,195]
d: [2,564,5,12,5,4,21,563,200]
d: [2,565,5,10,4,5,21,564,205]
d: [2,566,5,7,3,6,21,565,315]
d: [2,567,5,5,2,7,21,566,276]
d: [2,568,5,2,1,8,21,567,305]
d: [2,569,5,0,0,9,21,568,262]
e: 569
f: 569,524,508,494,480,443,377,353,311,301,296,288,272,270,242,84,85,86,87,88,89,90,91,92,93
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mssubfam__t36_mssubfam
% SZS output start Proof for mssubfam__t36_mssubfam
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : (m3_pboole(X3,X0,k1_mboolean(X0,X2)) => m3_pboole(k4_pboole(X0,X3,X1),X0,k1_mboolean(X0,X2))))) [input t36_mssubfam]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : (m3_pboole(X3,X0,k1_mboolean(X0,X2)) => m3_pboole(k4_pboole(X0,X3,X1),X0,k1_mboolean(X0,X2))))) [negated conjecture 1]
3. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (m3_pboole(X2,X0,X1) <=> r2_pboole(X0,X2,X1)))) [input d18_pboole]
6. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k4_pboole(X0,X1,X2),X0) & v1_funct_1(k4_pboole(X0,X1,X2)) & v4_relat_1(k4_pboole(X0,X1,X2),X0) & v1_relat_1(k4_pboole(X0,X1,X2)))) [input dt_k4_pboole]
7. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (m3_pboole(X2,X0,X1) => (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)))) [input dt_m3_pboole]
9. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k1_mboolean(X0,X1),X0) & v1_funct_1(k1_mboolean(X0,X1)) & v4_relat_1(k1_mboolean(X0,X1),X0) & v2_relat_1(k1_mboolean(X0,X1)) & v1_relat_1(k1_mboolean(X0,X1)))) [input fc1_mboolean]
21. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,X3))))) [input t13_pboole]
23. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k4_pboole(X0,X1,X2),X1))) [input t56_pboole]
24. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : (m3_pboole(X3,X0,k1_mboolean(X0,X2)) => m3_pboole(k4_pboole(X0,X3,X1),X0,k1_mboolean(X0,X2))))) [flattening 2]
25. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k4_pboole(X0,X1,X2),X1))) [flattening 23]
26. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (m3_pboole(X2,X0,X1) <=> r2_pboole(X0,X2,X1)))) [flattening 3]
28. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,X3))))) [flattening 21]
29. ? [X0,X1] : (? [X2] : (? [X3] : (~m3_pboole(k4_pboole(X0,X3,X1),X0,k1_mboolean(X0,X2)) & m3_pboole(X3,X0,k1_mboolean(X0,X2))) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 24]
30. ? [X0,X1] : (? [X2] : (? [X3] : (~m3_pboole(k4_pboole(X0,X3,X1),X0,k1_mboolean(X0,X2)) & m3_pboole(X3,X0,k1_mboolean(X0,X2))) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 29]
31. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k4_pboole(X0,X1,X2),X1) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 25]
32. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k4_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 31]
33. ! [X0,X1] : (! [X2] : ((m3_pboole(X2,X0,X1) <=> r2_pboole(X0,X2,X1)) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 26]
34. ! [X0,X1] : (! [X2] : ((m3_pboole(X2,X0,X1) <=> r2_pboole(X0,X2,X1)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 33]
37. ! [X0,X1] : (! [X2] : (! [X3] : ((r2_pboole(X0,X1,X3) | (~r2_pboole(X0,X2,X3) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 28]
38. ! [X0,X1] : (! [X2] : (! [X3] : (r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X2,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 37]
45. ! [X0,X1] : ((v1_partfun1(k1_mboolean(X0,X1),X0) & v1_funct_1(k1_mboolean(X0,X1)) & v4_relat_1(k1_mboolean(X0,X1),X0) & v2_relat_1(k1_mboolean(X0,X1)) & v1_relat_1(k1_mboolean(X0,X1))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 9]
46. ! [X0,X1] : ((v1_partfun1(k1_mboolean(X0,X1),X0) & v1_funct_1(k1_mboolean(X0,X1)) & v4_relat_1(k1_mboolean(X0,X1),X0) & v2_relat_1(k1_mboolean(X0,X1)) & v1_relat_1(k1_mboolean(X0,X1))) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 45]
49. ! [X0,X1] : (! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) | ~m3_pboole(X2,X0,X1)) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 7]
50. ! [X0,X1] : (! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) | ~m3_pboole(X2,X0,X1)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 49]
57. ! [X0,X1,X2] : ((v1_partfun1(k4_pboole(X0,X1,X2),X0) & v1_funct_1(k4_pboole(X0,X1,X2)) & v4_relat_1(k4_pboole(X0,X1,X2),X0) & v1_relat_1(k4_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 6]
58. ! [X0,X1,X2] : ((v1_partfun1(k4_pboole(X0,X1,X2),X0) & v1_funct_1(k4_pboole(X0,X1,X2)) & v4_relat_1(k4_pboole(X0,X1,X2),X0) & v1_relat_1(k4_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 57]
59. ? [X0,X1] : (? [X2] : (? [X3] : (~m3_pboole(k4_pboole(X0,X3,X1),X0,k1_mboolean(X0,X2)) & m3_pboole(X3,X0,k1_mboolean(X0,X2))) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (~m3_pboole(k4_pboole(sK0,X3,sK1),sK0,k1_mboolean(sK0,X2)) & m3_pboole(X3,sK0,k1_mboolean(sK0,X2))) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
60. ? [X2] : (? [X3] : (~m3_pboole(k4_pboole(sK0,X3,sK1),sK0,k1_mboolean(sK0,X2)) & m3_pboole(X3,sK0,k1_mboolean(sK0,X2))) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (? [X3] : (~m3_pboole(k4_pboole(sK0,X3,sK1),sK0,k1_mboolean(sK0,sK2)) & m3_pboole(X3,sK0,k1_mboolean(sK0,sK2))) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
61. ? [X3] : (~m3_pboole(k4_pboole(sK0,X3,sK1),sK0,k1_mboolean(sK0,sK2)) & m3_pboole(X3,sK0,k1_mboolean(sK0,sK2))) => (~m3_pboole(k4_pboole(sK0,sK3,sK1),sK0,k1_mboolean(sK0,sK2)) & m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2))) [choice axiom]
62. ((~m3_pboole(k4_pboole(sK0,sK3,sK1),sK0,k1_mboolean(sK0,sK2)) & m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2))) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 30,61,60,59]
63. ! [X0,X1] : (! [X2] : (((m3_pboole(X2,X0,X1) | ~r2_pboole(X0,X2,X1)) & (r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 34]
84. v1_relat_1(sK1) [cnf transformation 62] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
85. v4_relat_1(sK1,sK0) [cnf transformation 62] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
86. v1_funct_1(sK1) [cnf transformation 62] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
87. v1_partfun1(sK1,sK0) [cnf transformation 62] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
88. v1_relat_1(sK2) [cnf transformation 62] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
89. v4_relat_1(sK2,sK0) [cnf transformation 62] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
90. v1_funct_1(sK2) [cnf transformation 62] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
91. v1_partfun1(sK2,sK0) [cnf transformation 62] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
92. m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2)) [cnf transformation 62] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
93. ~m3_pboole(k4_pboole(sK0,sK3,sK1),sK0,k1_mboolean(sK0,sK2)) [cnf transformation 62] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
94. r2_pboole(X0,k4_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 32] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
95. r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 63] {a:0,w:28,wCS:11200,sine:2,thAx:0,allAx:1,thDist:-1}
96. ~r2_pboole(X0,X2,X1) | m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 63] {a:0,w:28,wCS:11200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
98. ~r2_pboole(X0,X2,X3) | r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 38] {a:0,w:42,wCS:16800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
125. v1_relat_1(k1_mboolean(X0,X1)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 46] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
127. v4_relat_1(k1_mboolean(X0,X1),X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 46] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
128. v1_funct_1(k1_mboolean(X0,X1)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 46] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
129. v1_partfun1(k1_mboolean(X0,X1),X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 46] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
134. ~m3_pboole(X2,X0,X1) | v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 50] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
135. ~m3_pboole(X2,X0,X1) | v4_relat_1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 50] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
136. ~m3_pboole(X2,X0,X1) | v1_funct_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 50] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
137. ~m3_pboole(X2,X0,X1) | v1_partfun1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 50] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
146. v1_relat_1(k4_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 58] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
147. v4_relat_1(k4_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 58] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
148. v1_funct_1(k4_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 58] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
149. v1_partfun1(k4_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 58] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
158. 1 <=> m3_pboole(k4_pboole(sK0,sK3,sK1),sK0,k1_mboolean(sK0,sK2)) [avatar definition]
160. ~m3_pboole(k4_pboole(sK0,sK3,sK1),sK0,k1_mboolean(sK0,sK2)) <- (~1) [avatar component clause 158] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
161. ~1 [avatar split clause 93,158]
163. 2 <=> m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2)) [avatar definition]
165. m3_pboole(sK3,sK0,k1_mboolean(sK0,sK2)) <- (2) [avatar component clause 163] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
166. 2 [avatar split clause 92,163]
168. 3 <=> v1_partfun1(sK2,sK0) [avatar definition]
170. v1_partfun1(sK2,sK0) <- (3) [avatar component clause 168] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
171. 3 [avatar split clause 91,168]
173. 4 <=> v1_funct_1(sK2) [avatar definition]
175. v1_funct_1(sK2) <- (4) [avatar component clause 173] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
176. 4 [avatar split clause 90,173]
178. 5 <=> v4_relat_1(sK2,sK0) [avatar definition]
180. v4_relat_1(sK2,sK0) <- (5) [avatar component clause 178] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
181. 5 [avatar split clause 89,178]
183. 6 <=> v1_relat_1(sK2) [avatar definition]
185. v1_relat_1(sK2) <- (6) [avatar component clause 183] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
186. 6 [avatar split clause 88,183]
188. 7 <=> v1_partfun1(sK1,sK0) [avatar definition]
190. v1_partfun1(sK1,sK0) <- (7) [avatar component clause 188] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
191. 7 [avatar split clause 87,188]
193. 8 <=> v1_funct_1(sK1) [avatar definition]
195. v1_funct_1(sK1) <- (8) [avatar component clause 193] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
196. 8 [avatar split clause 86,193]
198. 9 <=> v4_relat_1(sK1,sK0) [avatar definition]
200. v4_relat_1(sK1,sK0) <- (9) [avatar component clause 198] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
201. 9 [avatar split clause 85,198]
203. 10 <=> v1_relat_1(sK1) [avatar definition]
205. v1_relat_1(sK1) <- (10) [avatar component clause 203] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
206. 10 [avatar split clause 84,203]
242. v1_relat_1(sK3) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (2) [resolution 134,165] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
244. 18 <=> v1_relat_1(k1_mboolean(sK0,sK2)) [avatar definition]
245. v1_relat_1(k1_mboolean(sK0,sK2)) <- (18) [avatar component clause 244] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
246. ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (~18) [avatar component clause 244] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
248. 19 <=> v4_relat_1(k1_mboolean(sK0,sK2),sK0) [avatar definition]
249. v4_relat_1(k1_mboolean(sK0,sK2),sK0) <- (19) [avatar component clause 248] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
250. ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) <- (~19) [avatar component clause 248] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
252. 20 <=> v1_funct_1(k1_mboolean(sK0,sK2)) [avatar definition]
253. v1_funct_1(k1_mboolean(sK0,sK2)) <- (20) [avatar component clause 252] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
254. ~v1_funct_1(k1_mboolean(sK0,sK2)) <- (~20) [avatar component clause 252] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
256. 21 <=> v1_partfun1(k1_mboolean(sK0,sK2),sK0) [avatar definition]
257. v1_partfun1(k1_mboolean(sK0,sK2),sK0) <- (21) [avatar component clause 256] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
258. ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) <- (~21) [avatar component clause 256] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
260. 22 <=> v1_relat_1(sK3) [avatar definition]
262. v1_relat_1(sK3) <- (22) [avatar component clause 260] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
263. ~18 | ~19 | ~20 | ~21 | 22 | ~2 [avatar split clause 242,163,260,256,252,248,244]
265. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~18) [resolution 246,125] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
267. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (3, ~18) [subsumption resolution 265,170] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
268. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (3, 4, ~18) [subsumption resolution 267,175] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
269. ~v1_relat_1(sK2) <- (3, 4, 5, ~18) [subsumption resolution 268,180] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
270. $false <- (3, 4, 5, 6, ~18) [subsumption resolution 269,185] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
271. ~3 | ~4 | ~5 | ~6 | 18 [avatar contradiction clause 270]
272. v1_funct_1(sK3) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (2) [resolution 136,165] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
274. 23 <=> v1_funct_1(sK3) [avatar definition]
276. v1_funct_1(sK3) <- (23) [avatar component clause 274] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
277. ~18 | ~19 | ~20 | ~21 | 23 | ~2 [avatar split clause 272,163,274,256,252,248,244]
283. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~19) [resolution 250,127] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
285. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (3, ~19) [subsumption resolution 283,170] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
286. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (3, 4, ~19) [subsumption resolution 285,175] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
287. ~v1_relat_1(sK2) <- (3, 4, 5, ~19) [subsumption resolution 286,180] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
288. $false <- (3, 4, 5, 6, ~19) [subsumption resolution 287,185] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
289. ~3 | ~4 | ~5 | ~6 | 19 [avatar contradiction clause 288]
291. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~20) [resolution 254,128] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
293. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (3, ~20) [subsumption resolution 291,170] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
294. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (3, 4, ~20) [subsumption resolution 293,175] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
295. ~v1_relat_1(sK2) <- (3, 4, 5, ~20) [subsumption resolution 294,180] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
296. $false <- (3, 4, 5, 6, ~20) [subsumption resolution 295,185] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
297. ~3 | ~4 | ~5 | ~6 | 20 [avatar contradiction clause 296]
298. v4_relat_1(sK3,sK0) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (2) [resolution 135,165] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
301. v4_relat_1(sK3,sK0) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) <- (2, 18) [subsumption resolution 298,245] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
303. 24 <=> v4_relat_1(sK3,sK0) [avatar definition]
305. v4_relat_1(sK3,sK0) <- (24) [avatar component clause 303] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
306. ~19 | ~20 | ~21 | 24 | ~2 | ~18 [avatar split clause 301,244,163,303,256,252,248]
307. v1_partfun1(sK3,sK0) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (2) [resolution 137,165] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
310. v1_partfun1(sK3,sK0) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (2, 20) [subsumption resolution 307,253] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
311. v1_partfun1(sK3,sK0) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) <- (2, 18, 20) [subsumption resolution 310,245] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
313. 25 <=> v1_partfun1(sK3,sK0) [avatar definition]
315. v1_partfun1(sK3,sK0) <- (25) [avatar component clause 313] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
316. ~19 | ~21 | 25 | ~2 | ~18 | ~20 [avatar split clause 311,252,244,163,313,256,248]
348. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~21) [resolution 258,129] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
350. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (3, ~21) [subsumption resolution 348,170] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
351. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (3, 4, ~21) [subsumption resolution 350,175] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
352. ~v1_relat_1(sK2) <- (3, 4, 5, ~21) [subsumption resolution 351,180] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
353. $false <- (3, 4, 5, 6, ~21) [subsumption resolution 352,185] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
354. ~3 | ~4 | ~5 | ~6 | 21 [avatar contradiction clause 353]
359. r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [subsumption resolution 95,134] {a:0,w:26,wCS:10400,sine:1,thAx:0,allAx:1,thDist:-1}
360. r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v4_relat_1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [subsumption resolution 359,136] {a:0,w:24,wCS:9600,sine:1,thAx:0,allAx:1,thDist:-1}
361. r2_pboole(X0,X2,X1) | ~m3_pboole(X2,X0,X1) | ~v1_partfun1(X2,X0) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [subsumption resolution 360,135] {a:0,w:21,wCS:8400,sine:1,thAx:0,allAx:1,thDist:-1}
362. ~m3_pboole(X2,X0,X1) | r2_pboole(X0,X2,X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [subsumption resolution 361,137] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
364. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (2) [resolution 362,165] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
374. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (2, 21) [subsumption resolution 364,257] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
375. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (2, 20, 21) [subsumption resolution 374,253] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
376. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (2, 19, 20, 21) [subsumption resolution 375,249] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
377. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) <- (2, 18, 19, 20, 21) [subsumption resolution 376,245] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
379. 26 <=> r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) [avatar definition]
381. r2_pboole(sK0,sK3,k1_mboolean(sK0,sK2)) <- (26) [avatar component clause 379] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
382. 26 | ~2 | ~18 | ~19 | ~20 | ~21 [avatar split clause 377,256,252,248,244,163,379]
398. r2_pboole(sK0,X8,k1_mboolean(sK0,sK2)) | ~r2_pboole(sK0,X8,sK3) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X8,sK0) | ~v1_funct_1(X8) | ~v4_relat_1(X8,sK0) | ~v1_relat_1(X8) <- (26) [resolution 98,381] {a:2,w:48,wCS:4800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
405. r2_pboole(sK0,X8,k1_mboolean(sK0,sK2)) | ~r2_pboole(sK0,X8,sK3) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X8,sK0) | ~v1_funct_1(X8) | ~v4_relat_1(X8,sK0) | ~v1_relat_1(X8) <- (21, 26) [subsumption resolution 398,257] {a:2,w:43,wCS:4300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
406. r2_pboole(sK0,X8,k1_mboolean(sK0,sK2)) | ~r2_pboole(sK0,X8,sK3) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X8,sK0) | ~v1_funct_1(X8) | ~v4_relat_1(X8,sK0) | ~v1_relat_1(X8) <- (20, 21, 26) [subsumption resolution 405,253] {a:2,w:39,wCS:3900,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
407. r2_pboole(sK0,X8,k1_mboolean(sK0,sK2)) | ~r2_pboole(sK0,X8,sK3) | ~v1_relat_1(k1_mboolean(sK0,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X8,sK0) | ~v1_funct_1(X8) | ~v4_relat_1(X8,sK0) | ~v1_relat_1(X8) <- (19, 20, 21, 26) [subsumption resolution 406,249] {a:2,w:34,wCS:3400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
408. r2_pboole(sK0,X8,k1_mboolean(sK0,sK2)) | ~r2_pboole(sK0,X8,sK3) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X8,sK0) | ~v1_funct_1(X8) | ~v4_relat_1(X8,sK0) | ~v1_relat_1(X8) <- (18, 19, 20, 21, 26) [subsumption resolution 407,245] {a:2,w:30,wCS:3000,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
409. r2_pboole(sK0,X8,k1_mboolean(sK0,sK2)) | ~r2_pboole(sK0,X8,sK3) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X8,sK0) | ~v1_funct_1(X8) | ~v4_relat_1(X8,sK0) | ~v1_relat_1(X8) <- (18, 19, 20, 21, 25, 26) [subsumption resolution 408,315] {a:2,w:27,wCS:2700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
410. r2_pboole(sK0,X8,k1_mboolean(sK0,sK2)) | ~r2_pboole(sK0,X8,sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(X8,sK0) | ~v1_funct_1(X8) | ~v4_relat_1(X8,sK0) | ~v1_relat_1(X8) <- (18, 19, 20, 21, 23, 25, 26) [subsumption resolution 409,276] {a:2,w:25,wCS:2500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
411. r2_pboole(sK0,X8,k1_mboolean(sK0,sK2)) | ~r2_pboole(sK0,X8,sK3) | ~v1_relat_1(sK3) | ~v1_partfun1(X8,sK0) | ~v1_funct_1(X8) | ~v4_relat_1(X8,sK0) | ~v1_relat_1(X8) <- (18, 19, 20, 21, 23, 24, 25, 26) [subsumption resolution 410,305] {a:2,w:22,wCS:2200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
412. r2_pboole(sK0,X8,k1_mboolean(sK0,sK2)) | ~r2_pboole(sK0,X8,sK3) | ~v1_partfun1(X8,sK0) | ~v1_funct_1(X8) | ~v4_relat_1(X8,sK0) | ~v1_relat_1(X8) <- (18, 19, 20, 21, 22, 23, 24, 25, 26) [subsumption resolution 411,262] {a:2,w:20,wCS:2000,nSel:1,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
414. ~r2_pboole(sK0,X2,sK3) | ~v1_partfun1(X2,sK0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,sK0) | ~v1_relat_1(X2) | m3_pboole(X2,sK0,k1_mboolean(sK0,sK2)) | ~v1_partfun1(X2,sK0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,sK0) | ~v1_relat_1(X2) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (18, 19, 20, 21, 22, 23, 24, 25, 26) [resolution 412,96] {a:3,w:48,wCS:4800,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
415. ~r2_pboole(sK0,X2,sK3) | ~v1_partfun1(X2,sK0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,sK0) | ~v1_relat_1(X2) | m3_pboole(X2,sK0,k1_mboolean(sK0,sK2)) | ~v1_partfun1(k1_mboolean(sK0,sK2),sK0) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (18, 19, 20, 21, 22, 23, 24, 25, 26) [duplicate literal removal 414] {a:3,w:38,wCS:3800,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
421. ~r2_pboole(sK0,X2,sK3) | ~v1_partfun1(X2,sK0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,sK0) | ~v1_relat_1(X2) | m3_pboole(X2,sK0,k1_mboolean(sK0,sK2)) | ~v1_funct_1(k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (18, 19, 20, 21, 22, 23, 24, 25, 26) [subsumption resolution 415,257] {a:3,w:33,wCS:3300,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
422. ~r2_pboole(sK0,X2,sK3) | ~v1_partfun1(X2,sK0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,sK0) | ~v1_relat_1(X2) | m3_pboole(X2,sK0,k1_mboolean(sK0,sK2)) | ~v4_relat_1(k1_mboolean(sK0,sK2),sK0) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (18, 19, 20, 21, 22, 23, 24, 25, 26) [subsumption resolution 421,253] {a:3,w:29,wCS:2900,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
423. ~r2_pboole(sK0,X2,sK3) | ~v1_partfun1(X2,sK0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,sK0) | ~v1_relat_1(X2) | m3_pboole(X2,sK0,k1_mboolean(sK0,sK2)) | ~v1_relat_1(k1_mboolean(sK0,sK2)) <- (18, 19, 20, 21, 22, 23, 24, 25, 26) [subsumption resolution 422,249] {a:3,w:24,wCS:2400,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
424. m3_pboole(X2,sK0,k1_mboolean(sK0,sK2)) | ~v1_partfun1(X2,sK0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,sK0) | ~v1_relat_1(X2) | ~r2_pboole(sK0,X2,sK3) <- (18, 19, 20, 21, 22, 23, 24, 25, 26) [subsumption resolution 423,245] {a:3,w:20,wCS:2000,nSel:1,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
443. ~v1_partfun1(k4_pboole(sK0,sK3,sK1),sK0) | ~v1_funct_1(k4_pboole(sK0,sK3,sK1)) | ~v4_relat_1(k4_pboole(sK0,sK3,sK1),sK0) | ~v1_relat_1(k4_pboole(sK0,sK3,sK1)) | ~r2_pboole(sK0,k4_pboole(sK0,sK3,sK1),sK3) <- (~1, 18, 19, 20, 21, 22, 23, 24, 25, 26) [resolution 424,160] {a:4,w:29,wCS:2900,goal:1,sine:2,thAx:0,allAx:11,thDist:-11}
450. 27 <=> r2_pboole(sK0,k4_pboole(sK0,sK3,sK1),sK3) [avatar definition]
452. ~r2_pboole(sK0,k4_pboole(sK0,sK3,sK1),sK3) <- (~27) [avatar component clause 450] {a:4,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
454. 28 <=> v1_relat_1(k4_pboole(sK0,sK3,sK1)) [avatar definition]
456. ~v1_relat_1(k4_pboole(sK0,sK3,sK1)) <- (~28) [avatar component clause 454] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
458. 29 <=> v4_relat_1(k4_pboole(sK0,sK3,sK1),sK0) [avatar definition]
460. ~v4_relat_1(k4_pboole(sK0,sK3,sK1),sK0) <- (~29) [avatar component clause 458] {a:4,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
462. 30 <=> v1_funct_1(k4_pboole(sK0,sK3,sK1)) [avatar definition]
464. ~v1_funct_1(k4_pboole(sK0,sK3,sK1)) <- (~30) [avatar component clause 462] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
466. 31 <=> v1_partfun1(k4_pboole(sK0,sK3,sK1),sK0) [avatar definition]
468. ~v1_partfun1(k4_pboole(sK0,sK3,sK1),sK0) <- (~31) [avatar component clause 466] {a:4,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
469. ~27 | ~28 | ~29 | ~30 | ~31 | 1 | ~18 | ~19 | ~20 | ~21 | ~22 | ~23 | ~24 | ~25 | ~26 [avatar split clause 443,379,313,303,274,260,256,252,248,244,158,466,462,458,454,450]
471. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~27) [resolution 452,94] {a:5,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
473. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, ~27) [subsumption resolution 471,190] {a:5,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
474. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, ~27) [subsumption resolution 473,195] {a:5,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
475. ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, ~27) [subsumption resolution 474,200] {a:5,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
476. ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, ~27) [subsumption resolution 475,205] {a:5,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
477. ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 25, ~27) [subsumption resolution 476,315] {a:5,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
478. ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 25, ~27) [subsumption resolution 477,276] {a:5,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
479. ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 24, 25, ~27) [subsumption resolution 478,305] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
480. $false <- (7, 8, 9, 10, 22, 23, 24, 25, ~27) [subsumption resolution 479,262] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
481. ~7 | ~8 | ~9 | ~10 | ~22 | ~23 | ~24 | ~25 | 27 [avatar contradiction clause 480]
485. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~28) [resolution 456,146] {a:5,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
487. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, ~28) [subsumption resolution 485,190] {a:5,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
488. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, ~28) [subsumption resolution 487,195] {a:5,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
489. ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, ~28) [subsumption resolution 488,200] {a:5,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
490. ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, ~28) [subsumption resolution 489,205] {a:5,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
491. ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 25, ~28) [subsumption resolution 490,315] {a:5,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
492. ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 25, ~28) [subsumption resolution 491,276] {a:5,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
493. ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 24, 25, ~28) [subsumption resolution 492,305] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
494. $false <- (7, 8, 9, 10, 22, 23, 24, 25, ~28) [subsumption resolution 493,262] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
495. ~7 | ~8 | ~9 | ~10 | ~22 | ~23 | ~24 | ~25 | 28 [avatar contradiction clause 494]
499. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~29) [resolution 460,147] {a:5,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
501. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, ~29) [subsumption resolution 499,190] {a:5,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
502. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, ~29) [subsumption resolution 501,195] {a:5,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
503. ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, ~29) [subsumption resolution 502,200] {a:5,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
504. ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, ~29) [subsumption resolution 503,205] {a:5,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
505. ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 25, ~29) [subsumption resolution 504,315] {a:5,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
506. ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 25, ~29) [subsumption resolution 505,276] {a:5,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
507. ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 24, 25, ~29) [subsumption resolution 506,305] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
508. $false <- (7, 8, 9, 10, 22, 23, 24, 25, ~29) [subsumption resolution 507,262] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
509. ~7 | ~8 | ~9 | ~10 | ~22 | ~23 | ~24 | ~25 | 29 [avatar contradiction clause 508]
515. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~30) [resolution 464,148] {a:5,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
517. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, ~30) [subsumption resolution 515,190] {a:5,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
518. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, ~30) [subsumption resolution 517,195] {a:5,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
519. ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, ~30) [subsumption resolution 518,200] {a:5,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
520. ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, ~30) [subsumption resolution 519,205] {a:5,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
521. ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 25, ~30) [subsumption resolution 520,315] {a:5,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
522. ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 25, ~30) [subsumption resolution 521,276] {a:5,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
523. ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 24, 25, ~30) [subsumption resolution 522,305] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
524. $false <- (7, 8, 9, 10, 22, 23, 24, 25, ~30) [subsumption resolution 523,262] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
525. ~7 | ~8 | ~9 | ~10 | ~22 | ~23 | ~24 | ~25 | 30 [avatar contradiction clause 524]
560. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~31) [resolution 468,149] {a:5,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
562. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, ~31) [subsumption resolution 560,190] {a:5,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
563. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, ~31) [subsumption resolution 562,195] {a:5,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
564. ~v1_relat_1(sK1) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, ~31) [subsumption resolution 563,200] {a:5,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
565. ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, ~31) [subsumption resolution 564,205] {a:5,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
566. ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 25, ~31) [subsumption resolution 565,315] {a:5,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
567. ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 25, ~31) [subsumption resolution 566,276] {a:5,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
568. ~v1_relat_1(sK3) <- (7, 8, 9, 10, 23, 24, 25, ~31) [subsumption resolution 567,305] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
569. $false <- (7, 8, 9, 10, 22, 23, 24, 25, ~31) [subsumption resolution 568,262] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
570. ~7 | ~8 | ~9 | ~10 | ~22 | ~23 | ~24 | ~25 | 31 [avatar contradiction clause 569]
571. $false [avatar sat refutation 161,166,171,176,181,186,191,196,201,206,263,271,277,289,297,306,316,354,382,469,481,495,509,525,570] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:564,thDist:-564}
% SZS output end Proof for mssubfam__t36_mssubfam
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% Introduced skolems: 14
% 
% >>> Saturation
% Initial clauses: 73
% Generated clauses: 361
% Active clauses: 104
% Passive clauses: 121
% Final active clauses: 96
% Final passive clauses: 12
% Main loop iterations started: 109
% 
% >>> Simplifying Inferences
% Duplicate literals: 112
% Fw subsumption resolutions: 147
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 22
% 
% >>> Generating Inferences
% Binary resolution: 61
% Unit resulting resolution: 9
% Forward superposition: 6
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 22
% Sat splitting refutations: 18
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 43
% SAT solver unit clauses: 17
% 
% Memory used [KB]: 5373
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% term sharing: 0.001 s
% non unit clause index maintenance: 0.001 s
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.007 s ( own 0.001 s ) 

