i: [1,146,0,3,1,1,0,0] t2_relset_1
a: [3,215,0,3,1,146]
s: 215
i: [1,145,0,6,1,1,0,0] t2_relset_1
a: [3,220,0,6,1,145]
s: 220
i: [1,148,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,225,0,2,1,148]
s: 225
i: [1,149,0,2,1,0,0,255] fc1_xboole_0
a: [3,230,0,2,1,149]
s: 230
i: [1,207,0,2,1,0,0,2] rc2_xboole_0
a: [3,240,0,2,1,207]
s: 240
i: [1,208,0,2,1,0,0,2] rc1_xboole_0
a: [3,245,0,2,1,208]
s: 245
i: [1,210,0,2,1,0,0,2] rc1_relat_1
a: [3,250,0,2,1,210]
s: 250
i: [1,209,0,2,1,0,0,2] rc1_relat_1
a: [3,255,0,2,1,209]
s: 255
i: [1,211,0,2,1,0,0,255] rc2_relat_1
a: [3,260,0,2,1,211]
s: 260
i: [1,147,0,11,3,1,0,0] t2_relset_1
s: 147
i: [1,151,0,3,1,0,0,255] fc2_xboole_0
s: 151
i: [1,152,0,3,1,0,0,1] fc1_subset_1
s: 152
i: [1,188,0,3,1,0,0,1] rc2_subset_1
s: 188
i: [1,191,0,3,1,0,0,255] reflexivity_r1_tarski
s: 191
i: [1,150,0,3,1,0,0,255] d2_xboole_0
a: [3,235,0,3,1,150]
s: 235
r: 225
i: [1,160,0,4,2,0,0,2] cc1_relat_1
s: 160
d: [2,263,1,2,1,1,40,160,230]
a: [3,269,1,2,1,263]
s: 269
i: [1,186,0,4,1,0,0,1] existence_m1_subset_1
s: 186
d: [2,265,1,2,1,1,40,160,245]
a: [3,274,1,2,1,265]
s: 274
i: [1,190,0,4,1,0,0,255] rc3_subset_1
s: 190
i: [1,192,0,4,1,0,0,255] fc3_xboole_0
s: 192
d: [2,264,1,3,1,0,40,160,188]
s: 264
i: [1,193,0,4,1,0,0,1] fc6_relat_1
s: 193
i: [1,203,0,4,1,0,0,255] rc3_relat_1
s: 203
i: [1,156,0,5,2,0,0,1] rc1_subset_1
s: 156
i: [1,161,0,5,2,0,0,255] t6_boole
s: 161
i: [1,180,0,5,2,0,0,1] t7_boole
s: 180
i: [1,187,0,5,1,0,0,1] rc2_subset_1
s: 187
d: [2,279,1,2,1,1,40,180,215]
a: [3,283,1,2,1,279]
s: 283
i: [1,189,0,5,1,0,0,255] rc3_subset_1
s: 189
i: [1,194,0,5,1,0,0,255] fc5_relat_1
s: 194
i: [1,204,0,5,1,0,0,255] rc3_relat_1
s: 204
i: [1,205,0,5,1,0,0,255] rc3_relat_1
s: 205
i: [1,158,0,6,2,0,0,255] rc4_subset_1
s: 158
i: [1,173,0,6,2,0,0,1] d1_relat_1
s: 173
i: [1,176,0,6,2,0,0,1] t1_subset
s: 176
d: [2,286,1,3,1,1,40,176,215]
a: [3,291,1,3,1,286]
s: 291
i: [1,196,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 196
d: [2,293,1,3,1,1,40,196,215]
a: [3,298,1,3,1,293]
s: 298
i: [1,155,0,7,2,0,0,1] rc1_subset_1
s: 155
i: [1,157,0,7,2,0,0,255] rc4_subset_1
s: 157
d: [2,287,1,6,2,0,40,176,173]
s: 287
i: [1,178,0,7,2,0,0,255] t3_subset
s: 178
d: [2,294,1,6,2,0,40,196,173]
s: 294
i: [1,179,0,7,3,0,0,1] t8_boole
s: 179
i: [1,195,0,7,1,0,0,255] commutativity_k2_tarski
s: 195
d: [2,276,1,7,3,0,44,161,161]
s: 276
i: [1,159,0,8,3,0,0,2] cc2_relat_1
s: 159
d: [2,303,1,6,2,1,40,159,220]
d: [2,310,1,2,1,1,21,303,193]
a: [3,314,1,2,1,310]
s: 314
i: [1,162,0,8,3,0,0,1] cc1_subset_1
s: 162
i: [1,174,0,8,2,0,0,1] d1_relat_1
s: 174
d: [2,317,1,6,2,1,40,162,220]
d: [2,324,1,4,1,2,21,317,283]
a: [3,328,1,4,1,324]
s: 328
i: [1,177,0,8,3,0,0,1] t2_subset
s: 177
d: [2,332,1,11,2,1,40,177,220]
d: [2,340,1,6,1,1,21,332,152]
a: [3,344,1,6,1,340]
s: 344
i: [1,197,0,8,3,0,0,1] fc10_subset_1
s: 197
i: [1,200,0,8,2,0,0,1] cc1_relset_1
s: 200
d: [2,350,2,6,1,1,40,344,196]
a: [3,356,2,6,1,350]
s: 356
i: [1,206,0,8,1,0,0,255] fc7_relat_1
s: 206
d: [2,309,1,5,2,0,40,159,189]
s: 309
i: [1,163,0,9,3,0,0,255] cc4_subset_1
s: 163
i: [1,164,0,9,4,0,0,255] cc8_relat_1
s: 164
d: [2,323,1,5,2,0,40,162,189]
s: 323
i: [1,167,0,9,4,0,0,255] cc7_relat_1
s: 167
d: [2,337,1,8,2,0,40,177,187]
d: [2,348,1,5,1,0,21,337,152]
s: 348
i: [1,182,0,9,3,0,0,1] t5_subset
s: 182
i: [1,201,0,9,2,0,0,255] cc2_relset_1
s: 201
d: [2,393,1,3,1,1,40,201,220]
a: [3,403,1,3,1,393]
s: 403
i: [1,202,0,9,2,0,0,255] cc2_relset_1
s: 202
d: [2,405,2,6,3,1,40,403,167]
d: [2,406,2,4,2,2,21,405,283]
d: [2,407,2,2,1,3,21,406,314]
a: [3,411,2,2,1,407]
s: 411
i: [1,175,0,10,1,0,0,255] d5_tarski
s: 175
i: [1,181,0,10,3,0,0,1] t4_subset
s: 181
d: [2,414,1,3,1,1,40,202,220]
a: [3,424,1,3,1,414]
s: 424
i: [1,183,0,10,2,0,0,1] t87_zfmisc_1
s: 183
d: [2,442,2,6,3,1,40,424,164]
d: [2,443,2,4,2,2,21,442,283]
d: [2,444,2,2,1,3,21,443,314]
a: [3,448,2,2,1,444]
s: 448
i: [1,184,0,10,2,0,0,1] t87_zfmisc_1
s: 184
i: [1,153,0,11,4,0,0,255] cc3_subset_1
s: 153
d: [2,435,1,8,2,1,40,181,220]
s: 435
i: [1,154,0,11,4,0,0,255] cc2_subset_1
s: 154
d: [2,462,2,12,3,1,40,435,177]
d: [2,463,2,8,2,2,21,462,328]
s: 463
i: [1,172,0,12,3,0,0,1] d1_relat_1
s: 172
i: [1,198,0,12,4,0,0,255] cc6_relat_1
s: 198
d: [2,477,1,12,3,1,40,198,220]
d: [2,484,1,8,2,1,21,477,193]
s: 484
i: [1,199,0,12,4,0,0,255] cc5_relat_1
s: 199
d: [2,486,1,12,3,1,40,199,220]
d: [2,493,1,8,2,1,21,486,193]
s: 493
i: [1,185,0,13,3,0,0,1] t87_zfmisc_1
s: 185
d: [2,469,1,16,5,0,44,147,172]
s: 469
d: [2,500,2,13,4,0,47,469]
a: [3,507,2,4,1,500]
s: 507
d: [2,433,1,4,1,0,44,192,175]
s: 433
d: [2,464,3,8,2,2,40,463,196]
s: 464
d: [2,338,1,8,2,0,40,177,189]
d: [2,349,1,5,1,0,21,338,152]
s: 349
d: [2,365,1,5,1,0,40,200,189]
s: 365
d: [2,514,4,10,2,2,40,464,463]
d: [2,515,4,5,1,2,18,514]
a: [3,519,4,5,1,515]
s: 519
d: [2,457,1,9,3,0,40,153,189]
d: [2,461,1,5,2,0,21,457,190]
s: 461
d: [2,467,3,8,2,2,40,463,184]
s: 467
d: [2,307,1,6,2,0,40,159,186]
s: 307
d: [2,321,1,6,2,0,40,162,186]
s: 321
d: [2,468,3,8,2,2,40,463,183]
s: 468
d: [2,336,1,6,2,0,40,177,186]
s: 336
d: [2,525,4,12,4,2,44,467,172]
s: 525
d: [2,531,5,0,0,5,52,314,215,215,507,525]
e: 531
r: 507
d: [2,363,1,6,1,0,40,200,186]
s: 363
d: [2,374,1,6,2,0,40,163,187]
s: 374
a: [3,506,2,4,1,500]
s: 506
d: [2,376,1,10,3,0,40,164,205]
d: [2,377,1,6,2,0,21,376,203]
s: 377
a: [3,511,2,4,1,500]
s: 511
d: [2,379,1,10,3,0,40,167,204]
d: [2,380,1,6,2,0,21,379,203]
s: 380
d: [2,390,1,6,2,0,40,182,187]
s: 390
d: [2,538,3,4,1,1,40,506,196]
a: [3,544,3,4,1,538]
s: 544
d: [2,391,1,6,2,0,40,182,189]
s: 391
d: [2,539,3,4,1,1,40,506,176]
a: [3,549,3,4,1,539]
s: 549
d: [2,398,1,6,1,0,40,201,187]
s: 398
d: [2,399,1,6,1,0,40,201,189]
s: 399
d: [2,527,4,12,4,2,44,468,172]
s: 527
d: [2,561,5,0,0,5,52,314,215,215,511,527]
e: 561
r: 511
d: [2,419,1,6,1,0,40,202,187]
s: 419
a: [3,510,2,4,1,500]
s: 510
d: [2,420,1,6,1,0,40,202,189]
s: 420
d: [2,456,1,9,3,0,40,153,187]
d: [2,460,1,6,2,0,21,456,188]
s: 460
d: [2,569,3,4,1,1,40,510,196]
a: [3,575,3,4,1,569]
s: 575
d: [2,302,1,7,3,0,40,159,178]
s: 302
d: [2,570,3,4,1,1,40,510,176]
a: [3,580,3,4,1,570]
s: 580
d: [2,304,1,7,3,0,40,159,155]
s: 304
d: [2,305,1,7,3,0,40,159,157]
s: 305
a: [3,503,2,5,2,500]
s: 503
d: [2,587,3,0,0,3,52,314,215,503]
e: 587
d: [2,588,3,2,1,2,40,503,215]
d: [2,591,3,0,0,3,21,588,314]
e: 591
d: [2,589,3,7,2,3,40,503,463]
d: [2,593,3,3,1,3,21,589,193]
d: [2,594,3,0,0,4,21,593,215]
e: 594
f: 591,561,531,500,324,310,279,145,146
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relset_1__t2_relset_1
% SZS output start Proof for relset_1__t2_relset_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))) => ~(! [X4] : ! [X5] : ~(r2_hidden(X5,X2) & r2_hidden(X4,X1) & k4_tarski(X4,X5) = X0) & r2_hidden(X0,X3))) [input t2_relset_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))) => ~(! [X4] : ! [X5] : ~(r2_hidden(X5,X2) & r2_hidden(X4,X1) & k4_tarski(X4,X5) = X0) & r2_hidden(X0,X3))) [negated conjecture 1]
6. ! [X0] : (v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_xboole_0(X1))) [input cc1_subset_1]
7. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
19. ! [X0] : (v1_relat_1(X0) <=> ! [X1] : ~(! [X2] : ! [X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) [input d1_relat_1]
37. ! [X0,X1] : v1_relat_1(k2_zfmisc_1(X0,X1)) [input fc6_relat_1]
50. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
52. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
55. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
56. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [input t87_zfmisc_1]
58. ~! [X0,X1,X2,X3] : (m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))) => ~(! [X4,X5] : ~(r2_hidden(X5,X2) & r2_hidden(X4,X1) & k4_tarski(X4,X5) = X0) & r2_hidden(X0,X3))) [flattening 2]
59. ! [X0] : (v1_relat_1(X0) <=> ! [X1] : ~(! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) [flattening 19]
62. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 50]
65. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 55]
66. ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 52]
68. ! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [flattening 56]
76. ? [X0,X1,X2,X3] : ((! [X4,X5] : (~r2_hidden(X5,X2) | ~r2_hidden(X4,X1) | k4_tarski(X4,X5) != X0) & r2_hidden(X0,X3)) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))) [ennf transformation 58]
77. ? [X0,X1,X2,X3] : (! [X4,X5] : (~r2_hidden(X5,X2) | ~r2_hidden(X4,X1) | k4_tarski(X4,X5) != X0) & r2_hidden(X0,X3) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))) [flattening 76]
84. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 7]
87. ! [X0] : (! [X1] : (v1_xboole_0(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_xboole_0(X0)) [ennf transformation 6]
97. ! [X0] : (v1_relat_1(X0) <=> ! [X1] : (? [X2,X3] : k4_tarski(X2,X3) = X1 | ~r2_hidden(X1,X0))) [ennf transformation 59]
99. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 62]
100. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 99]
103. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 65]
104. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))) [ennf transformation 66]
105. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [flattening 104]
116. ? [X0,X1,X2,X3] : (! [X4,X5] : (~r2_hidden(X5,X2) | ~r2_hidden(X4,X1) | k4_tarski(X4,X5) != X0) & r2_hidden(X0,X3) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))) => (! [X5,X4] : (~r2_hidden(X5,sK2) | ~r2_hidden(X4,sK1) | k4_tarski(X4,X5) != sK0) & r2_hidden(sK0,sK3) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK1,sK2)))) [choice axiom]
117. ! [X4,X5] : (~r2_hidden(X5,sK2) | ~r2_hidden(X4,sK1) | k4_tarski(X4,X5) != sK0) & r2_hidden(sK0,sK3) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK1,sK2))) [skolemisation 77,116]
122. ! [X0] : ((v1_relat_1(X0) | ? [X1] : (! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) & (! [X1] : (? [X2,X3] : k4_tarski(X2,X3) = X1 | ~r2_hidden(X1,X0)) | ~v1_relat_1(X0))) [nnf transformation 97]
123. ! [X0] : ((v1_relat_1(X0) | ? [X1] : (! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) & (! [X4] : (? [X5,X6] : k4_tarski(X5,X6) = X4 | ~r2_hidden(X4,X0)) | ~v1_relat_1(X0))) [rectify 122]
124. ! [X0] : (? [X1] : (! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0)) => (! [X3,X2] : k4_tarski(X2,X3) != sK6(X0) & r2_hidden(sK6(X0),X0))) [choice axiom]
125. ! [X4] : (? [X5,X6] : k4_tarski(X5,X6) = X4 => k4_tarski(sK7(X4),sK8(X4)) = X4) [choice axiom]
126. ! [X0] : ((v1_relat_1(X0) | (! [X2,X3] : k4_tarski(X2,X3) != sK6(X0) & r2_hidden(sK6(X0),X0))) & (! [X4] : (k4_tarski(sK7(X4),sK8(X4)) = X4 | ~r2_hidden(X4,X0)) | ~v1_relat_1(X0))) [skolemisation 123,125,124]
127. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | (~r2_hidden(X1,X3) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [nnf transformation 68]
128. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [flattening 127]
145. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK1,sK2))) [cnf transformation 117] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
146. r2_hidden(sK0,sK3) [cnf transformation 117] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
147. k4_tarski(X4,X5) != sK0 | ~r2_hidden(X4,sK1) | ~r2_hidden(X5,sK2) [cnf transformation 117] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
159. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 84] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
162. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 87] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
172. k4_tarski(sK7(X4),sK8(X4)) = X4 | ~r2_hidden(X4,X0) | ~v1_relat_1(X0) [cnf transformation 126] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
177. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 100] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
180. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 103] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
181. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1) [cnf transformation 105] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
183. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X0,X2) [cnf transformation 128] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
184. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X1,X3) [cnf transformation 128] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
193. v1_relat_1(k2_zfmisc_1(X0,X1)) [cnf transformation 37] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
213. 1 <=> r2_hidden(sK0,sK3) [avatar definition]
215. r2_hidden(sK0,sK3) <- (1) [avatar component clause 213] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
216. 1 [avatar split clause 146,213]
218. 2 <=> m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK1,sK2))) [avatar definition]
220. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK1,sK2))) <- (2) [avatar component clause 218] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
221. 2 [avatar split clause 145,218]
279. ~v1_xboole_0(sK3) <- (1) [resolution 180,215] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
281. 13 <=> v1_xboole_0(sK3) [avatar definition]
283. ~v1_xboole_0(sK3) <- (~13) [avatar component clause 281] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
284. ~13 | ~1 [avatar split clause 279,213,281]
303. v1_relat_1(sK3) | ~v1_relat_1(k2_zfmisc_1(sK1,sK2)) <- (2) [resolution 159,220] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
310. v1_relat_1(sK3) <- (2) [subsumption resolution 303,193] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
312. 16 <=> v1_relat_1(sK3) [avatar definition]
314. v1_relat_1(sK3) <- (16) [avatar component clause 312] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
315. 16 | ~2 [avatar split clause 310,218,312]
317. v1_xboole_0(sK3) | ~v1_xboole_0(k2_zfmisc_1(sK1,sK2)) <- (2) [resolution 162,220] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
324. ~v1_xboole_0(k2_zfmisc_1(sK1,sK2)) <- (2, ~13) [subsumption resolution 317,283] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
326. 17 <=> v1_xboole_0(k2_zfmisc_1(sK1,sK2)) [avatar definition]
328. ~v1_xboole_0(k2_zfmisc_1(sK1,sK2)) <- (~17) [avatar component clause 326] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
329. ~17 | ~2 | 13 [avatar split clause 324,281,218,326]
435. m1_subset_1(X3,k2_zfmisc_1(sK1,sK2)) | ~r2_hidden(X3,sK3) <- (2) [resolution 181,220] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
462. ~r2_hidden(X0,sK3) | v1_xboole_0(k2_zfmisc_1(sK1,sK2)) | r2_hidden(X0,k2_zfmisc_1(sK1,sK2)) <- (2) [resolution 435,177] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
463. r2_hidden(X0,k2_zfmisc_1(sK1,sK2)) | ~r2_hidden(X0,sK3) <- (2, ~17) [subsumption resolution 462,328] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
467. ~r2_hidden(k4_tarski(X3,X4),sK3) | r2_hidden(X4,sK2) <- (2, ~17) [resolution 463,184] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
468. ~r2_hidden(k4_tarski(X5,X6),sK3) | r2_hidden(X5,sK1) <- (2, ~17) [resolution 463,183] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
469. sK0 != X0 | ~r2_hidden(sK7(X0),sK1) | ~r2_hidden(sK8(X0),sK2) | ~r2_hidden(X0,X1) | ~v1_relat_1(X1) [superposition 147,172] {a:1,w:16,wCS:1600,nSel:1,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
500. ~r2_hidden(sK7(sK0),sK1) | ~r2_hidden(sK8(sK0),sK2) | ~r2_hidden(sK0,X0) | ~v1_relat_1(X0) [equality resolution 469] {a:2,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
502. 24 <=> ! [X0] : (~r2_hidden(sK0,X0) | ~v1_relat_1(X0)) [avatar definition]
503. ~r2_hidden(sK0,X0) | ~v1_relat_1(X0) <- (24) [avatar component clause 502] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
505. 25 <=> r2_hidden(sK8(sK0),sK2) [avatar definition]
507. ~r2_hidden(sK8(sK0),sK2) <- (~25) [avatar component clause 505] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
509. 26 <=> r2_hidden(sK7(sK0),sK1) [avatar definition]
511. ~r2_hidden(sK7(sK0),sK1) <- (~26) [avatar component clause 509] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
512. 24 | ~25 | ~26 [avatar split clause 500,509,505,502]
525. r2_hidden(sK8(X0),sK2) | ~r2_hidden(X0,sK3) | ~r2_hidden(X0,X1) | ~v1_relat_1(X1) <- (2, ~17) [superposition 467,172] {a:4,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
527. r2_hidden(sK7(X0),sK1) | ~r2_hidden(X0,sK3) | ~r2_hidden(X0,X1) | ~v1_relat_1(X1) <- (2, ~17) [superposition 468,172] {a:4,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
531. $false <- (1, 2, 16, ~17, ~25) [unit resulting resolution 314,215,215,507,525] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:20,thDist:-20}
536. ~1 | ~2 | ~16 | 17 | 25 [avatar contradiction clause 531]
561. $false <- (1, 2, 16, ~17, ~26) [unit resulting resolution 314,215,215,511,527] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:20,thDist:-20}
566. ~1 | ~2 | ~16 | 17 | 26 [avatar contradiction clause 561]
588. ~v1_relat_1(sK3) <- (1, 24) [resolution 503,215] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
591. $false <- (1, 16, 24) [subsumption resolution 588,314] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
592. ~1 | ~16 | ~24 [avatar contradiction clause 591]
596. $false [avatar sat refutation 216,221,284,315,329,512,536,566,592] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:217,thDist:-217}
% SZS output end Proof for relset_1__t2_relset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 56
% >>> Preprocessing
% Introduced skolems: 17
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 67
% Generated clauses: 331
% Active clauses: 120
% Passive clauses: 202
% Final active clauses: 117
% Final passive clauses: 80
% Main loop iterations started: 122
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 26
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 19
% Forward subsumptions: 43
% 
% >>> Generating Inferences
% Binary resolution: 167
% Unit resulting resolution: 3
% Forward superposition: 10
% Backward superposition: 14
% Self superposition: 2
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 3
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 36
% SAT solver unit clauses: 10
% SAT solver binary clauses: 14
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% term sharing: 0.001 s
% forward subsumption resolution: 0.001 s
% resolution: 0.001 s ( own 0.0000 s ) 
% unit resulting resolution: 0.001 s
% other: 0.004 s ( own 0.0000 s ) 

