i: [1,138,0,2,1,1,0,0] t1_real
a: [3,199,0,2,1,138]
s: 199
i: [1,137,0,2,1,1,0,0] t1_real
a: [3,204,0,2,1,137]
s: 204
i: [1,135,0,2,1,1,0,0] t1_real
a: [3,214,0,2,1,135]
s: 214
i: [1,134,0,2,1,1,0,0] t1_real
a: [3,219,0,2,1,134]
s: 219
i: [1,136,0,3,1,1,0,0] t1_real
a: [3,209,0,3,1,136]
s: 209
i: [1,139,0,2,1,0,0,3] fc1_xboole_0
a: [3,224,0,2,1,139]
s: 224
i: [1,145,0,2,1,0,0,4] fc6_ordinal1
a: [3,249,0,2,1,145]
s: 249
i: [1,144,0,2,1,0,0,4] fc6_ordinal1
a: [3,254,0,2,1,144]
s: 254
i: [1,179,0,2,1,0,0,3] rc2_xboole_0
a: [3,266,0,2,1,179]
s: 266
i: [1,180,0,2,1,0,0,5] rc4_ordinal1
a: [3,271,0,2,1,180]
s: 271
i: [1,181,0,2,1,0,0,5] rc1_ordinal1
a: [3,276,0,2,1,181]
s: 276
i: [1,182,0,2,1,0,0,1] rc1_xreal_0
a: [3,281,0,2,1,182]
s: 281
i: [1,183,0,2,1,0,0,3] rc1_xboole_0
a: [3,286,0,2,1,183]
s: 286
i: [1,187,0,2,1,0,0,255] rc2_ordinal1
a: [3,291,0,2,1,187]
s: 291
i: [1,186,0,2,1,0,0,255] rc2_ordinal1
a: [3,296,0,2,1,186]
s: 296
i: [1,185,0,2,1,0,0,255] rc2_ordinal1
a: [3,301,0,2,1,185]
s: 301
i: [1,184,0,2,1,0,0,255] rc2_ordinal1
a: [3,306,0,2,1,184]
s: 306
i: [1,189,0,2,1,0,0,5] rc5_ordinal1
a: [3,311,0,2,1,189]
s: 311
i: [1,188,0,2,1,0,0,5] rc5_ordinal1
a: [3,316,0,2,1,188]
s: 316
i: [1,192,0,2,1,0,0,1] rc2_xreal_0
a: [3,321,0,2,1,192]
s: 321
i: [1,191,0,2,1,0,0,1] rc2_xreal_0
a: [3,326,0,2,1,191]
s: 326
i: [1,190,0,2,1,0,0,1] rc2_xreal_0
a: [3,331,0,2,1,190]
s: 331
i: [1,195,0,2,1,0,0,2] rc4_xreal_0
a: [3,336,0,2,1,195]
s: 336
i: [1,194,0,2,1,0,0,2] rc4_xreal_0
a: [3,341,0,2,1,194]
s: 341
i: [1,193,0,2,1,0,0,2] rc4_xreal_0
a: [3,346,0,2,1,193]
s: 346
i: [1,171,0,3,1,0,0,255] reflexivity_r1_tarski
s: 171
i: [1,140,0,3,1,0,0,2] redefinition_k6_numbers
a: [3,229,0,3,1,140]
s: 229
r: 224
d: [2,348,0,2,1,2,23,224,229]
a: [3,352,0,2,1,348]
s: 352
i: [1,141,0,3,1,0,0,3] redefinition_k5_numbers
a: [3,234,0,3,1,141]
s: 234
i: [1,149,0,4,2,0,0,5] cc6_ordinal1
s: 149
i: [1,150,0,4,2,0,0,5] cc2_xreal_0
s: 150
d: [2,354,1,2,1,1,40,149,271]
a: [3,359,1,2,1,354]
s: 359
i: [1,151,0,4,2,0,0,255] cc1_ordinal1
s: 151
d: [2,355,1,2,1,1,40,149,311]
a: [3,364,1,2,1,355]
s: 364
i: [1,152,0,4,2,0,0,255] cc1_ordinal1
s: 152
i: [1,154,0,4,2,0,0,2] cc4_xreal_0
s: 154
d: [2,428,1,2,1,1,40,154,219]
a: [3,436,1,2,1,428]
s: 436
i: [1,156,0,4,2,0,0,255] cc9_ordinal1
s: 156
d: [2,429,1,2,1,1,40,154,214]
a: [3,441,1,2,1,429]
s: 441
i: [1,157,0,4,2,0,0,5] cc7_ordinal1
s: 157
i: [1,158,0,4,2,0,0,5] cc3_ordinal1
s: 158
d: [2,366,1,2,1,1,40,150,271]
a: [3,371,1,2,1,366]
s: 371
i: [1,170,0,4,1,0,0,4] existence_m1_subset_1
s: 170
d: [2,367,1,2,1,1,40,150,311]
a: [3,376,1,2,1,367]
s: 376
i: [1,142,0,4,1,0,0,3] dt_k5_numbers
a: [3,239,0,4,1,142]
s: 239
i: [1,143,0,4,1,0,0,2] dt_k6_numbers
a: [3,244,0,4,1,143]
s: 244
d: [2,378,1,2,1,1,40,151,249]
a: [3,385,1,2,1,378]
s: 385
d: [2,514,0,4,1,2,22,239,234]
a: [3,518,0,4,1,514]
s: 518
d: [2,380,1,2,1,1,40,151,276]
a: [3,395,1,2,1,380]
s: 395
d: [2,520,0,4,1,2,22,244,234]
a: [3,524,0,4,1,520]
s: 524
i: [1,159,0,5,2,0,0,3] t6_boole
s: 159
d: [2,403,1,2,1,1,40,152,249]
a: [3,411,1,2,1,403]
s: 411
i: [1,160,0,5,2,0,0,3] cc1_xreal_0
s: 160
d: [2,405,1,2,1,1,40,152,276]
a: [3,421,1,2,1,405]
s: 421
i: [1,161,0,5,2,0,0,4] cc8_ordinal1
s: 161
i: [1,167,0,5,2,0,0,255] t7_boole
s: 167
d: [2,430,1,2,1,1,40,154,281]
a: [3,446,1,2,1,430]
s: 446
i: [1,177,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,261,0,5,2,177]
s: 261
d: [2,448,1,2,1,1,40,156,346]
a: [3,454,1,2,1,448]
s: 454
d: [2,526,0,5,2,1,22,159,229]
s: 526
i: [1,162,0,6,3,0,0,255] cc2_ordinal1
s: 162
d: [2,449,1,2,1,1,40,156,286]
a: [3,459,1,2,1,449]
s: 459
i: [1,163,0,6,2,0,0,255] t1_subset
s: 163
d: [2,450,1,2,1,1,40,156,352]
a: [3,464,1,2,1,450]
s: 464
i: [1,172,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 172
i: [1,147,0,7,3,0,0,1] d6_xxreal_0
s: 147
d: [2,466,1,2,1,1,40,157,346]
a: [3,472,1,2,1,466]
s: 472
i: [1,148,0,7,3,0,0,1] d6_xxreal_0
s: 148
d: [2,467,1,2,1,1,40,157,286]
a: [3,477,1,2,1,467]
s: 477
i: [1,153,0,7,3,0,0,5] cc5_ordinal1
s: 153
i: [1,165,0,7,2,0,0,255] t3_subset
s: 165
d: [2,468,1,2,1,1,40,157,352]
a: [3,482,1,2,1,468]
s: 482
i: [1,166,0,7,3,0,0,3] t8_boole
s: 166
d: [2,484,1,2,1,1,40,158,346]
a: [3,490,1,2,1,484]
s: 490
i: [1,146,0,8,3,0,0,255] cc10_ordinal1
s: 146
i: [1,164,0,8,3,0,0,255] t2_subset
s: 164
d: [2,485,1,2,1,1,40,158,286]
a: [3,495,1,2,1,485]
s: 495
i: [1,169,0,9,3,0,0,255] t5_subset
s: 169
d: [2,486,1,2,1,1,40,158,352]
a: [3,500,1,2,1,486]
s: 500
i: [1,168,0,10,3,0,0,255] t4_subset
s: 168
i: [1,178,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 178
d: [2,546,1,6,3,1,40,147,261]
d: [2,548,1,4,2,1,18,546]
a: [3,552,1,2,1,548]
s: 552
i: [1,176,0,14,4,0,0,3] existence_m2_subset_1
s: 176
d: [2,597,1,4,2,1,40,146,518]
a: [3,602,1,2,1,597]
s: 602
i: [1,155,0,15,6,0,0,1] l2_real
s: 155
i: [1,173,0,15,5,0,0,3] dt_m2_subset_1
s: 173
d: [2,668,1,12,5,1,40,155,209]
d: [2,673,1,10,4,2,21,668,214]
d: [2,674,1,8,3,3,21,673,219]
s: 674
i: [1,174,0,15,5,0,0,3] redefinition_m2_subset_1
s: 174
d: [2,633,1,5,2,1,40,169,518]
a: [3,640,1,2,1,633]
s: 640
d: [2,706,2,0,0,2,52,602,640,156]
e: 706
d: [2,709,2,2,1,1,40,640,156]
d: [2,711,2,0,0,2,21,709,602]
e: 711
r: 640
i: [1,175,0,15,5,0,0,3] redefinition_m2_subset_1
s: 175
r: 602
a: [3,606,1,2,1,597]
s: 606
s: 640
r: 640
a: [3,641,1,2,1,633]
s: 641
d: [2,693,2,12,5,3,40,674,178]
d: [2,694,2,10,4,4,21,693,436]
d: [2,695,2,8,3,4,21,694,154]
s: 695
d: [2,527,1,3,1,0,40,160,170]
a: [3,531,1,3,1,527]
s: 531
d: [2,533,1,3,1,0,40,161,170]
a: [3,537,1,3,1,533]
s: 537
d: [2,735,3,9,4,4,40,695,147]
d: [2,740,3,7,3,5,21,735,204]
d: [2,741,3,5,2,5,21,740,436]
a: [3,745,3,3,1,741]
s: 745
a: [3,637,1,3,1,633]
s: 637
d: [2,773,2,0,0,2,52,170,254,637,164]
e: 773
r: 637
d: [2,767,4,5,2,2,44,745,526]
s: 767
d: [2,682,1,11,4,1,40,173,524]
d: [2,685,1,7,3,2,21,682,518]
d: [2,686,1,5,2,3,21,685,254]
a: [3,690,1,3,1,686]
s: 690
d: [2,696,1,11,4,1,40,174,524]
d: [2,699,1,7,3,2,21,696,518]
d: [2,700,1,5,2,3,21,699,254]
a: [3,704,1,3,1,700]
s: 704
d: [2,379,2,2,1,1,40,151,359]
a: [3,390,2,2,1,379]
s: 390
d: [2,609,1,7,2,1,40,164,518]
a: [3,614,1,4,1,609]
s: 614
d: [2,397,2,2,1,1,40,364,151]
a: [3,401,2,2,1,397]
s: 401
a: [3,617,1,3,1,609]
s: 617
d: [2,572,1,5,2,0,40,153,170]
s: 572
d: [2,404,2,2,1,1,40,152,359]
a: [3,416,2,2,1,404]
s: 416
d: [2,540,1,6,2,2,44,524,526]
s: 540
d: [2,407,2,2,1,1,40,152,364]
a: [3,426,2,2,1,407]
s: 426
d: [2,596,1,6,2,0,40,146,170]
s: 596
d: [2,608,1,6,2,0,40,164,170]
s: 608
d: [2,502,2,2,1,1,40,371,154]
a: [3,506,2,2,1,502]
s: 506
d: [2,657,1,6,2,1,40,168,518]
s: 657
d: [2,508,2,2,1,1,40,376,154]
a: [3,512,2,2,1,508]
s: 512
d: [2,539,1,7,3,1,44,526,526]
s: 539
d: [2,598,1,7,3,0,40,146,165]
s: 598
d: [2,564,2,2,1,1,40,477,150]
a: [3,569,2,2,1,564]
s: 569
d: [2,632,1,7,2,0,40,169,170]
s: 632
d: [2,575,2,2,1,1,40,482,150]
a: [3,581,2,2,1,575]
s: 581
d: [2,837,3,0,0,2,52,552,581,154]
e: 837
d: [2,838,3,2,1,1,40,581,154]
d: [2,841,3,0,0,2,21,838,552]
e: 841
r: 552
a: [3,551,1,2,1,548]
s: 551
a: [3,556,1,2,1,548]
s: 556
d: [2,584,2,2,1,1,40,490,152]
a: [3,589,2,2,1,584]
s: 589
d: [2,659,1,11,5,0,40,178,147]
d: [2,662,1,9,4,0,18,659]
a: [3,665,1,7,3,662]
s: 665
d: [2,585,2,2,1,1,40,490,151]
a: [3,594,2,2,1,585]
s: 594
d: [2,714,1,11,4,1,40,175,518]
d: [2,716,1,9,3,2,21,714,254]
a: [3,719,1,7,2,716]
s: 719
d: [2,574,1,8,3,0,40,165,153]
s: 574
d: [2,620,2,2,1,1,40,495,152]
a: [3,625,2,2,1,620]
s: 625
d: [2,634,1,8,3,0,40,169,165]
s: 634
d: [2,621,2,2,1,1,40,495,151]
a: [3,630,2,2,1,621]
s: 630
d: [2,656,1,8,2,0,40,168,170]
s: 656
d: [2,547,1,9,4,1,44,147,526]
s: 547
d: [2,863,2,6,3,2,40,547,209]
d: [2,874,2,4,2,3,21,863,204]
d: [2,875,2,2,1,4,21,874,436]
a: [3,879,2,2,1,875]
s: 879
d: [2,562,1,9,4,1,44,148,526]
s: 562
d: [2,888,2,11,5,4,40,562,674]
d: [2,892,2,9,4,4,21,888,154]
a: [3,896,2,2,1,892]
s: 896
d: [2,658,1,9,3,0,40,168,165]
s: 658
d: [2,610,1,10,3,0,40,164,165]
s: 610
d: [2,643,2,2,1,1,40,500,152]
a: [3,649,2,2,1,643]
s: 649
d: [2,670,1,16,7,0,40,155,148]
d: [2,675,1,14,6,0,21,670,154]
a: [3,678,1,12,5,675]
s: 678
d: [2,907,2,9,4,2,40,678,209]
d: [2,918,2,7,3,3,21,907,199]
d: [2,919,2,5,2,4,21,918,219]
d: [2,920,2,3,1,5,21,919,214]
a: [3,924,2,3,1,920]
s: 924
d: [2,940,3,0,0,3,52,436,204,924,147]
e: 940
d: [2,941,3,0,0,5,52,352,436,204,924,547]
e: 941
d: [2,942,3,4,2,1,40,924,147]
d: [2,949,3,2,1,2,21,942,204]
d: [2,950,3,0,0,3,21,949,436]
e: 950
d: [2,944,3,6,3,2,40,924,547]
d: [2,952,3,4,2,3,21,944,204]
d: [2,953,3,2,1,4,21,952,436]
d: [2,954,3,0,0,5,21,953,352]
e: 954
d: [2,777,2,2,1,1,40,690,160]
f: 954,920,777,773,686,675,633,520,514,428,348,144,143,142,141,140,139,134,135,136,137,138
% Refutation found. Thanks to Tanya!
% SZS status Theorem for real__t1_real
% SZS output start Proof for real__t1_real
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [negated conjecture 1]
6. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xreal_0(X0)) [input cc1_xreal_0]
12. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
19. ! [X0] : (v1_xxreal_0(X0) => (v2_xxreal_0(X0) <=> ~r1_xxreal_0(X0,k6_numbers))) [input d6_xxreal_0]
24. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [input dt_k5_numbers]
25. m2_subset_1(k6_numbers,k1_numbers,k5_numbers) [input dt_k6_numbers]
27. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ! [X2] : (m2_subset_1(X2,X0,X1) => m1_subset_1(X2,X0))) [input dt_m2_subset_1]
28. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
30. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
31. v3_ordinal1(k4_ordinal1) & ~v1_xboole_0(k4_ordinal1) [input fc6_ordinal1]
32. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ! [X2] : (v1_xreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input l2_real]
42. k4_ordinal1 = k5_numbers [input redefinition_k5_numbers]
43. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
48. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
51. ! [X0] : ! [X1] : ! [X2] : ~(v1_xboole_0(X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) [input t5_subset]
52. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
56. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 48]
61. ! [X0,X1,X2] : ~(v1_xboole_0(X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) [flattening 51]
68. ? [X0] : (? [X1] : ((~v2_xxreal_0(X1) & (v2_xxreal_0(X0) & r1_xxreal_0(X0,X1))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
69. ? [X0] : (? [X1] : (~v2_xxreal_0(X1) & v2_xxreal_0(X0) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 68]
71. ! [X0] : ((v2_xxreal_0(X0) <=> ~r1_xxreal_0(X0,k6_numbers)) | ~v1_xxreal_0(X0)) [ennf transformation 19]
76. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 12]
77. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 32]
78. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 77]
82. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 52]
83. ! [X0] : (v1_xreal_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 6]
88. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 56]
89. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 88]
95. ! [X0,X1,X2] : (~v1_xboole_0(X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [ennf transformation 61]
97. ! [X0,X1] : (! [X2] : (m1_subset_1(X2,X0) | ~m2_subset_1(X2,X0,X1)) | (~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 27]
98. ! [X0,X1] : (! [X2] : (m1_subset_1(X2,X0) | ~m2_subset_1(X2,X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 97]
107. ? [X0] : (? [X1] : (~v2_xxreal_0(X1) & v2_xxreal_0(X0) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~v2_xxreal_0(X1) & v2_xxreal_0(sK0) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
108. ? [X1] : (~v2_xxreal_0(X1) & v2_xxreal_0(sK0) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) => (~v2_xxreal_0(sK1) & v2_xxreal_0(sK0) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) [choice axiom]
109. (~v2_xxreal_0(sK1) & v2_xxreal_0(sK0) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 69,108,107]
110. ! [X0] : (((v2_xxreal_0(X0) | r1_xxreal_0(X0,k6_numbers)) & (~r1_xxreal_0(X0,k6_numbers) | ~v2_xxreal_0(X0))) | ~v1_xxreal_0(X0)) [nnf transformation 71]
111. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK2(X0),X0)) [choice axiom]
112. ! [X0] : m1_subset_1(sK2(X0),X0) [skolemisation 28,111]
134. v1_xreal_0(sK0) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
135. v1_xreal_0(sK1) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
136. r1_xxreal_0(sK0,sK1) [cnf transformation 109] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
137. v2_xxreal_0(sK0) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
138. ~v2_xxreal_0(sK1) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
139. v1_xboole_0(k1_xboole_0) [cnf transformation 30] {a:0,w:2,wCS:800,sine:3,thAx:0,allAx:1,thDist:-1}
140. k6_numbers = k1_xboole_0 [cnf transformation 43] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
141. k4_ordinal1 = k5_numbers [cnf transformation 42] {a:0,w:3,wCS:1200,sine:3,thAx:0,allAx:1,thDist:-1}
142. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [cnf transformation 24] {a:0,w:4,wCS:1600,sine:3,thAx:0,allAx:1,thDist:-1}
143. m2_subset_1(k6_numbers,k1_numbers,k5_numbers) [cnf transformation 25] {a:0,w:4,wCS:1600,sine:2,thAx:0,allAx:1,thDist:-1}
144. ~v1_xboole_0(k4_ordinal1) [cnf transformation 31] {a:0,w:2,wCS:800,sine:4,thAx:0,allAx:1,thDist:-1}
147. ~r1_xxreal_0(X0,k6_numbers) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 110] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
148. r1_xxreal_0(X0,k6_numbers) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 110] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
154. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 76] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
155. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 78] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
159. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 82] {a:0,w:5,wCS:2000,sine:3,thAx:0,allAx:1,thDist:-1}
160. ~m1_subset_1(X0,k1_numbers) | v1_xreal_0(X0) [cnf transformation 83] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
164. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 89] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
169. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~v1_xboole_0(X2) | ~r2_hidden(X0,X1) [cnf transformation 95] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
170. m1_subset_1(sK2(X0),X0) [cnf transformation 112] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:2,thDist:-2}
173. ~m2_subset_1(X2,X0,X1) | m1_subset_1(X2,X0) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 98] {a:0,w:15,wCS:6000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
197. 1 <=> v2_xxreal_0(sK1) [avatar definition]
199. ~v2_xxreal_0(sK1) <- (~1) [avatar component clause 197] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
200. ~1 [avatar split clause 138,197]
202. 2 <=> v2_xxreal_0(sK0) [avatar definition]
204. v2_xxreal_0(sK0) <- (2) [avatar component clause 202] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
205. 2 [avatar split clause 137,202]
207. 3 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
209. r1_xxreal_0(sK0,sK1) <- (3) [avatar component clause 207] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
210. 3 [avatar split clause 136,207]
212. 4 <=> v1_xreal_0(sK1) [avatar definition]
214. v1_xreal_0(sK1) <- (4) [avatar component clause 212] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
215. 4 [avatar split clause 135,212]
217. 5 <=> v1_xreal_0(sK0) [avatar definition]
219. v1_xreal_0(sK0) <- (5) [avatar component clause 217] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
220. 5 [avatar split clause 134,217]
222. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
224. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 222] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
225. 6 [avatar split clause 139,222]
227. 7 <=> k6_numbers = k1_xboole_0 [avatar definition]
229. k6_numbers = k1_xboole_0 <- (7) [avatar component clause 227] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
230. 7 [avatar split clause 140,227]
232. 8 <=> k4_ordinal1 = k5_numbers [avatar definition]
234. k4_ordinal1 = k5_numbers <- (8) [avatar component clause 232] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
235. 8 [avatar split clause 141,232]
237. 9 <=> m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [avatar definition]
239. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) <- (9) [avatar component clause 237] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
240. 9 [avatar split clause 142,237]
242. 10 <=> m2_subset_1(k6_numbers,k1_numbers,k5_numbers) [avatar definition]
244. m2_subset_1(k6_numbers,k1_numbers,k5_numbers) <- (10) [avatar component clause 242] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
245. 10 [avatar split clause 143,242]
252. 12 <=> v1_xboole_0(k4_ordinal1) [avatar definition]
254. ~v1_xboole_0(k4_ordinal1) <- (~12) [avatar component clause 252] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
255. ~12 [avatar split clause 144,252]
348. v1_xboole_0(k6_numbers) <- (6, 7) [backward demodulation 224,229] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
350. 32 <=> v1_xboole_0(k6_numbers) [avatar definition]
352. v1_xboole_0(k6_numbers) <- (32) [avatar component clause 350] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
353. 32 | ~6 | ~7 [avatar split clause 348,227,222,350]
428. v1_xxreal_0(sK0) <- (5) [resolution 154,219] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
434. 45 <=> v1_xxreal_0(sK0) [avatar definition]
436. v1_xxreal_0(sK0) <- (45) [avatar component clause 434] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
437. 45 | ~5 [avatar split clause 428,217,434]
514. m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) <- (8, 9) [forward demodulation 239,234] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
516. 59 <=> m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) [avatar definition]
518. m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) <- (59) [avatar component clause 516] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
519. 59 | ~8 | ~9 [avatar split clause 514,237,232,516]
520. m2_subset_1(k6_numbers,k1_numbers,k4_ordinal1) <- (8, 10) [forward demodulation 244,234] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
522. 60 <=> m2_subset_1(k6_numbers,k1_numbers,k4_ordinal1) [avatar definition]
524. m2_subset_1(k6_numbers,k1_numbers,k4_ordinal1) <- (60) [avatar component clause 522] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
525. 60 | ~8 | ~10 [avatar split clause 520,242,232,522]
526. k6_numbers = X0 | ~v1_xboole_0(X0) <- (7) [forward demodulation 159,229] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
547. ~r1_xxreal_0(X1,X0) | ~v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v1_xboole_0(X0) <- (7) [superposition 147,526] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
579. 66 <=> v1_xreal_0(k6_numbers) [avatar definition]
633. ~v1_xboole_0(k1_numbers) | ~r2_hidden(X2,k4_ordinal1) <- (59) [resolution 169,518] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:2,thDist:-2}
636. 75 <=> ! [X2] : ~r2_hidden(X2,k4_ordinal1) [avatar definition]
637. ~r2_hidden(X2,k4_ordinal1) <- (75) [avatar component clause 636] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
639. 76 <=> v1_xboole_0(k1_numbers) [avatar definition]
642. 75 | ~76 | ~59 [avatar split clause 633,516,639,636]
670. r1_xxreal_0(X3,k6_numbers) | ~r1_xxreal_0(X3,X4) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X4) | ~v1_xreal_0(X3) | v2_xxreal_0(X4) | ~v1_xxreal_0(X4) [resolution 155,148] {a:1,w:16,wCS:6400,sine:1,thAx:0,allAx:2,thDist:-2}
675. r1_xxreal_0(X3,k6_numbers) | ~r1_xxreal_0(X3,X4) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X4) | ~v1_xreal_0(X3) | v2_xxreal_0(X4) [subsumption resolution 670,154] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
677. 80 <=> ! [X3,X4] : (r1_xxreal_0(X3,k6_numbers) | v2_xxreal_0(X4) | ~v1_xreal_0(X3) | ~v1_xreal_0(X4) | ~r1_xxreal_0(X3,X4)) [avatar definition]
678. ~r1_xxreal_0(X3,X4) | v2_xxreal_0(X4) | ~v1_xreal_0(X3) | ~v1_xreal_0(X4) | r1_xxreal_0(X3,k6_numbers) <- (80) [avatar component clause 677] {a:1,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
679. ~66 | 80 [avatar split clause 675,677,579]
682. m1_subset_1(k6_numbers,k1_numbers) | ~m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) | v1_xboole_0(k4_ordinal1) | v1_xboole_0(k1_numbers) <- (60) [resolution 173,524] {a:1,w:11,wCS:4400,sine:3,thAx:0,allAx:2,thDist:-2}
685. m1_subset_1(k6_numbers,k1_numbers) | v1_xboole_0(k4_ordinal1) | v1_xboole_0(k1_numbers) <- (59, 60) [subsumption resolution 682,518] {a:1,w:7,wCS:2800,sine:3,thAx:0,allAx:2,thDist:-2}
686. m1_subset_1(k6_numbers,k1_numbers) | v1_xboole_0(k1_numbers) <- (~12, 59, 60) [subsumption resolution 685,254] {a:1,w:5,wCS:2000,sine:3,thAx:0,allAx:2,thDist:-2}
688. 81 <=> m1_subset_1(k6_numbers,k1_numbers) [avatar definition]
690. m1_subset_1(k6_numbers,k1_numbers) <- (81) [avatar component clause 688] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
691. 76 | 81 | 12 | ~59 | ~60 [avatar split clause 686,522,516,252,688,639]
773. $false <- (~12, 75) [unit resulting resolution 170,254,637,164] {a:2,w:0,sine:4,thAx:0,allAx:6,thDist:-6}
774. 12 | ~75 [avatar contradiction clause 773]
777. v1_xreal_0(k6_numbers) <- (81) [resolution 690,160] {a:2,w:2,wCS:800,sine:3,thAx:0,allAx:3,thDist:-3}
781. 66 | ~81 [avatar split clause 777,688,579]
907. v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | ~v1_xreal_0(sK1) | r1_xxreal_0(sK0,k6_numbers) <- (3, 80) [resolution 678,209] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
918. ~v1_xreal_0(sK0) | ~v1_xreal_0(sK1) | r1_xxreal_0(sK0,k6_numbers) <- (~1, 3, 80) [subsumption resolution 907,199] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
919. ~v1_xreal_0(sK1) | r1_xxreal_0(sK0,k6_numbers) <- (~1, 3, 5, 80) [subsumption resolution 918,219] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
920. r1_xxreal_0(sK0,k6_numbers) <- (~1, 3, 4, 5, 80) [subsumption resolution 919,214] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
922. 100 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
924. r1_xxreal_0(sK0,k6_numbers) <- (100) [avatar component clause 922] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
925. 100 | 1 | ~3 | ~4 | ~5 | ~80 [avatar split clause 920,677,217,212,207,197,922]
944. ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) | ~v1_xboole_0(k6_numbers) <- (7, 100) [resolution 924,547] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
952. ~v1_xxreal_0(sK0) | ~v1_xboole_0(k6_numbers) <- (2, 7, 100) [subsumption resolution 944,204] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
953. ~v1_xboole_0(k6_numbers) <- (2, 7, 45, 100) [subsumption resolution 952,436] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
954. $false <- (2, 7, 32, 45, 100) [subsumption resolution 953,352] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
955. ~2 | ~7 | ~32 | ~45 | ~100 [avatar contradiction clause 954]
958. $false [avatar sat refutation 200,205,210,215,220,225,230,235,240,245,255,353,437,519,525,642,679,691,774,781,925,955] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:595,thDist:-595}
% SZS output end Proof for real__t1_real
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 53
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 62
% Generated clauses: 419
% Active clauses: 135
% Passive clauses: 182
% Final active clauses: 129
% Final passive clauses: 41
% Main loop iterations started: 139
% 
% >>> Simplifying Inferences
% Duplicate literals: 24
% Fw subsumption resolutions: 55
% Fw demodulations: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Forward subsumptions: 38
% 
% >>> Generating Inferences
% Binary resolution: 155
% Unit resulting resolution: 5
% Factoring: 1
% Forward superposition: 15
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 15
% Split components: 30
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 114
% SAT solver unit clauses: 31
% SAT solver binary clauses: 56
% 
% Memory used [KB]: 5628
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s ( own 0.0000 s ) 
% sine selection: 0.001 s
% backward subsumption: 0.001 s
% other: 0.005 s ( own 0.003 s ) 

