i: [1,19,0,3,1,1,0,0] t11_xtuple_0
a: [3,33,0,3,1,19]
s: 33
i: [1,20,0,5,1,1,0,0] t11_xtuple_0
a: [3,28,0,5,1,20]
s: 28
i: [1,24,0,3,1,0,0,1] reflexivity_r1_tarski
s: 24
i: [1,21,0,6,1,0,0,255] d15_xtuple_0
s: 21
i: [1,22,0,8,2,0,0,255] t8_xtuple_0
s: 22
i: [1,23,0,8,2,0,0,255] t9_xtuple_0
s: 23
d: [2,35,1,9,2,0,44,23,21]
s: 35
d: [2,36,1,9,2,0,44,23,21]
s: 36
d: [2,37,2,10,2,0,44,35,21]
s: 37
d: [2,38,2,10,2,0,44,36,21]
s: 38
d: [2,39,3,8,2,0,40,37,22]
s: 39
d: [2,41,4,0,0,2,52,33,28,39]
e: 41
d: [2,42,4,3,1,1,40,39,28]
d: [2,44,4,0,0,2,21,42,33]
e: 44
f: 44,19,20
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t11_xtuple_0
% SZS output start Proof for xtuple_0__t11_xtuple_0
1. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k12_xtuple_0(X0),k12_xtuple_0(X1))) [input t11_xtuple_0]
2. ~! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k12_xtuple_0(X0),k12_xtuple_0(X1))) [negated conjecture 1]
3. ! [X0] : k12_xtuple_0(X0) = k10_xtuple_0(k9_xtuple_0(X0)) [input d15_xtuple_0]
8. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) [input t8_xtuple_0]
9. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1))) [input t9_xtuple_0]
10. ~! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k12_xtuple_0(X0),k12_xtuple_0(X1))) [flattening 2]
11. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) [flattening 8]
12. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1))) [flattening 9]
14. ? [X0,X1] : (~r1_tarski(k12_xtuple_0(X0),k12_xtuple_0(X1)) & r1_tarski(X0,X1)) [ennf transformation 10]
15. ! [X0,X1] : (r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 11]
16. ! [X0,X1] : (r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 12]
17. ? [X0,X1] : (~r1_tarski(k12_xtuple_0(X0),k12_xtuple_0(X1)) & r1_tarski(X0,X1)) => (~r1_tarski(k12_xtuple_0(sK0),k12_xtuple_0(sK1)) & r1_tarski(sK0,sK1)) [choice axiom]
18. ~r1_tarski(k12_xtuple_0(sK0),k12_xtuple_0(sK1)) & r1_tarski(sK0,sK1) [skolemisation 14,17]
19. r1_tarski(sK0,sK1) [cnf transformation 18] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
20. ~r1_tarski(k12_xtuple_0(sK0),k12_xtuple_0(sK1)) [cnf transformation 18] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
21. k12_xtuple_0(X0) = k10_xtuple_0(k9_xtuple_0(X0)) [cnf transformation 3] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
22. r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 15] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
23. r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 16] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
26. 1 <=> r1_tarski(k12_xtuple_0(sK0),k12_xtuple_0(sK1)) [avatar definition]
28. ~r1_tarski(k12_xtuple_0(sK0),k12_xtuple_0(sK1)) <- (~1) [avatar component clause 26] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
29. ~1 [avatar split clause 20,26]
31. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
33. r1_tarski(sK0,sK1) <- (2) [avatar component clause 31] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
34. 2 [avatar split clause 19,31]
35. r1_tarski(k12_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(k9_xtuple_0(X0),X1) [superposition 23,21] {a:1,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
37. ~r1_tarski(k9_xtuple_0(X1),k9_xtuple_0(X0)) | r1_tarski(k12_xtuple_0(X1),k12_xtuple_0(X0)) [superposition 35,21] {a:2,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
39. r1_tarski(k12_xtuple_0(X0),k12_xtuple_0(X1)) | ~r1_tarski(X0,X1) [resolution 37,22] {a:3,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:4,thDist:-4}
42. ~r1_tarski(sK0,sK1) <- (~1) [resolution 39,28] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
44. $false <- (~1, 2) [subsumption resolution 42,33] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
45. 1 | ~2 [avatar contradiction clause 44]
46. $false [avatar sat refutation 29,34,45] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:20,thDist:-20}
% SZS output end Proof for xtuple_0__t11_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 6
% Generated clauses: 17
% Active clauses: 10
% Passive clauses: 11
% Final active clauses: 10
% Main loop iterations started: 11
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 1
% 
% >>> Deletion Inferences
% Forward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 1
% Forward superposition: 4
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 6
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

