i: [1,48,0,2,1,1,0,0] t27_projpl_1
a: [3,106,0,2,1,48]
s: 106
i: [1,47,0,2,1,1,0,0] t27_projpl_1
a: [3,111,0,2,1,47]
s: 111
i: [1,46,0,2,1,1,0,0] t27_projpl_1
a: [3,116,0,2,1,46]
s: 116
i: [1,45,0,2,1,1,0,0] t27_projpl_1
a: [3,121,0,2,1,45]
s: 121
i: [1,44,0,2,1,1,0,0] t27_projpl_1
a: [3,126,0,2,1,44]
s: 126
i: [1,43,0,2,1,1,0,0] t27_projpl_1
a: [3,131,0,2,1,43]
s: 131
i: [1,52,0,4,1,1,0,0] t27_projpl_1
a: [3,86,0,4,1,52]
s: 86
i: [1,51,0,4,1,1,0,0] t27_projpl_1
a: [3,91,0,4,1,51]
s: 91
i: [1,50,0,4,1,1,0,0] t27_projpl_1
a: [3,96,0,4,1,50]
s: 96
i: [1,49,0,4,1,1,0,0] t27_projpl_1
a: [3,101,0,4,1,49]
s: 101
i: [1,54,0,5,1,1,0,0] t27_projpl_1
a: [3,76,0,5,1,54]
s: 76
i: [1,53,0,6,1,1,0,0] t27_projpl_1
a: [3,81,0,6,1,53]
s: 81
i: [1,67,0,2,1,0,0,255] rc2_xboole_0
a: [3,136,0,2,1,67]
s: 136
i: [1,68,0,2,1,0,0,255] rc1_xboole_0
a: [3,141,0,2,1,68]
s: 141
i: [1,69,0,2,1,0,0,1] existence_l1_incsp_1
a: [3,146,0,2,1,69]
s: 146
i: [1,66,0,4,1,0,0,1] existence_m1_subset_1
s: 66
i: [1,70,0,4,1,0,0,2] d5_zfmisc_1
s: 70
i: [1,71,0,4,1,0,0,2] d5_zfmisc_1
s: 71
i: [1,72,0,4,1,0,0,2] d5_zfmisc_1
s: 72
i: [1,55,0,5,2,0,0,255] dt_u1_incsp_1
s: 55
i: [1,65,0,23,7,0,0,1] d6_zfmisc_1
s: 65
d: [2,148,1,18,6,1,40,65,76]
a: [3,156,1,3,1,148]
s: 156
r: 81
r: 76
r: 91
d: [2,175,0,5,1,2,23,76,156]
a: [3,185,0,5,1,175]
s: 185
d: [2,174,0,6,1,2,23,81,156]
a: [3,180,0,6,1,174]
s: 180
i: [1,56,0,29,7,0,0,1] d6_projpl_1
s: 56
i: [1,57,0,29,7,0,0,1] d6_projpl_1
s: 57
d: [2,189,1,23,6,1,40,56,180]
d: [2,190,1,19,5,1,18,189]
d: [2,191,1,15,4,2,21,190,86]
d: [2,192,1,11,3,3,21,191,101]
d: [2,193,1,7,2,4,21,192,96]
d: [2,194,1,5,1,5,21,193,106]
a: [3,198,1,5,1,194]
s: 198
i: [1,58,0,29,7,0,0,1] d6_projpl_1
s: 58
d: [2,200,1,23,6,1,40,57,180]
d: [2,201,1,19,5,1,18,200]
d: [2,202,1,15,4,2,21,201,86]
d: [2,203,1,11,3,3,21,202,101]
d: [2,204,1,7,2,4,21,203,96]
d: [2,205,1,5,1,5,21,204,106]
a: [3,209,1,5,1,205]
s: 209
i: [1,59,0,29,7,0,0,1] d6_projpl_1
s: 59
i: [1,61,0,33,11,0,0,1] t26_projpl_1
s: 61
d: [2,233,2,0,0,9,52,131,126,121,116,111,106,101,96,71,198,101,61]
e: 233
d: [2,234,2,28,10,1,40,61,198]
d: [2,236,2,24,9,1,18,234]
d: [2,238,2,20,8,1,21,236,71]
d: [2,239,2,16,7,2,21,238,101]
d: [2,240,2,12,6,3,21,239,96]
d: [2,241,2,10,5,4,21,240,106]
d: [2,242,2,8,4,5,21,241,111]
d: [2,243,2,6,3,6,21,242,116]
d: [2,244,2,4,2,7,21,243,121]
d: [2,245,2,2,1,8,21,244,126]
d: [2,246,2,0,0,9,21,245,131]
e: 246
r: 156
r: 180
r: 198
r: 209
a: [3,155,1,3,1,148]
s: 155
s: 91
a: [3,160,1,3,1,148]
s: 160
r: 185
r: 86
s: 76
d: [2,262,0,5,1,2,23,185,160]
a: [3,267,0,5,1,262]
s: 267
d: [2,269,0,5,1,2,22,76,160]
a: [3,273,0,5,1,269]
s: 273
s: 81
r: 160
a: [3,152,1,3,1,148]
s: 152
r: 273
r: 267
r: 96
s: 86
s: 76
a: [3,159,1,3,1,148]
s: 159
s: 185
d: [2,285,0,5,1,2,23,273,152]
a: [3,291,0,5,1,285]
s: 291
d: [2,286,0,5,1,2,23,267,152]
a: [3,296,0,5,1,286]
s: 296
d: [2,298,0,5,1,2,22,76,152]
a: [3,302,0,5,1,298]
s: 302
d: [2,304,0,5,1,2,22,185,152]
a: [3,308,0,5,1,304]
s: 308
s: 81
r: 152
a: [3,151,1,3,1,148]
s: 151
s: 96
s: 76
a: [3,164,1,3,1,148]
s: 164
r: 76
r: 302
r: 291
r: 91
s: 185
s: 267
s: 273
d: [2,329,0,5,1,2,23,76,164]
a: [3,337,0,5,1,329]
s: 337
d: [2,330,0,5,1,2,23,302,164]
a: [3,342,0,5,1,330]
s: 342
d: [2,331,0,5,1,2,23,291,164]
a: [3,347,0,5,1,331]
s: 347
d: [2,353,0,5,1,2,22,273,164]
a: [3,357,0,5,1,353]
s: 357
s: 81
a: [3,179,0,6,1,174]
s: 179
a: [3,197,1,5,1,194]
s: 197
d: [2,367,0,6,1,2,22,81,164]
a: [3,371,0,6,1,367]
s: 371
a: [3,208,1,5,1,205]
s: 208
i: [1,60,0,44,10,0,0,1] d6_projpl_1
s: 60
d: [2,211,1,23,6,1,40,58,180]
d: [2,212,1,19,5,1,18,211]
d: [2,213,1,15,4,2,21,212,86]
d: [2,214,1,11,3,3,21,213,101]
d: [2,215,1,7,2,4,21,214,96]
d: [2,216,1,5,1,5,21,215,106]
a: [3,220,1,5,1,216]
s: 220
d: [2,423,2,0,0,9,52,131,126,121,116,111,106,101,86,101,71,220,61]
e: 423
d: [2,424,2,28,10,1,40,220,61]
d: [2,425,2,24,9,1,18,424]
d: [2,427,2,20,8,1,21,425,71]
d: [2,428,2,16,7,2,21,427,101]
d: [2,429,2,12,6,3,21,428,86]
d: [2,430,2,10,5,4,21,429,106]
d: [2,431,2,8,4,5,21,430,111]
d: [2,432,2,6,3,6,21,431,116]
d: [2,433,2,4,2,7,21,432,121]
d: [2,434,2,2,1,8,21,433,126]
d: [2,435,2,0,0,9,21,434,131]
e: 435
r: 220
a: [3,219,1,5,1,216]
s: 219
d: [2,222,1,23,6,1,40,59,180]
d: [2,223,1,19,5,1,18,222]
d: [2,224,1,15,4,2,21,223,86]
d: [2,225,1,11,3,3,21,224,101]
d: [2,226,1,7,2,4,21,225,96]
d: [2,227,1,5,1,5,21,226,106]
a: [3,231,1,5,1,227]
s: 231
d: [2,437,2,28,10,1,40,231,61]
d: [2,438,2,24,9,2,21,437,96]
d: [2,439,2,20,8,3,21,438,101]
d: [2,440,2,16,7,4,21,439,86]
d: [2,441,2,14,6,5,21,440,106]
d: [2,442,2,12,5,6,21,441,111]
d: [2,443,2,10,4,7,21,442,116]
d: [2,444,2,8,3,8,21,443,121]
d: [2,445,2,6,2,9,21,444,126]
d: [2,446,2,4,1,10,21,445,131]
a: [3,450,2,4,1,446]
s: 450
d: [2,374,1,23,6,1,40,371,58]
d: [2,379,1,19,5,1,18,374]
d: [2,386,1,15,4,2,21,379,86]
d: [2,387,1,11,3,3,21,386,96]
d: [2,388,1,7,2,4,21,387,101]
d: [2,389,1,5,1,5,21,388,106]
a: [3,393,1,5,1,389]
s: 393
r: 164
r: 371
r: 393
s: 91
a: [3,163,1,3,1,148]
s: 163
s: 76
a: [3,172,1,3,1,148]
s: 172
r: 76
r: 86
r: 179
r: 185
r: 208
r: 219
r: 231
r: 308
r: 337
r: 342
r: 450
r: 159
r: 172
s: 86
s: 76
s: 160
r: 76
r: 86
s: 185
a: [3,171,1,3,1,148]
s: 171
s: 273
s: 291
s: 450
d: [2,556,2,4,1,2,22,450,160]
d: [2,557,2,0,0,2,21,556,72]
e: 557
r: 151
r: 160
s: 86
s: 152
r: 96
r: 163
r: 197
r: 267
r: 273
r: 347
r: 357
r: 152
s: 96
s: 76
s: 151
s: 185
s: 160
r: 185
r: 76
r: 86
s: 267
s: 273
s: 163
s: 302
s: 171
s: 308
s: 337
s: 342
s: 347
s: 357
s: 81
s: 179
s: 197
d: [2,279,0,6,1,2,22,81,160]
a: [3,283,0,6,1,279]
s: 283
r: 160
a: [3,168,1,3,1,148]
s: 168
s: 86
s: 76
s: 171
s: 185
s: 302
s: 308
s: 337
s: 342
d: [2,655,0,5,1,2,22,76,168]
a: [3,659,0,5,1,655]
s: 659
d: [2,662,0,5,1,2,22,185,168]
a: [3,666,0,5,1,662]
s: 666
d: [2,668,0,5,1,2,22,302,168]
a: [3,672,0,5,1,668]
s: 672
d: [2,674,0,5,1,2,22,308,168]
a: [3,678,0,5,1,674]
s: 678
d: [2,680,0,5,1,2,22,337,168]
a: [3,684,0,5,1,680]
s: 684
d: [2,686,0,5,1,2,22,342,168]
a: [3,690,0,5,1,686]
s: 690
s: 208
r: 151
r: 168
r: 283
s: 86
s: 76
s: 152
r: 76
r: 96
r: 163
r: 197
r: 267
r: 273
r: 347
r: 357
r: 659
r: 666
r: 672
r: 678
r: 684
r: 690
s: 185
s: 159
s: 302
s: 308
s: 171
s: 337
d: [2,235,2,28,10,1,40,61,209]
d: [2,248,2,24,9,2,21,235,86]
d: [2,249,2,20,8,3,21,248,101]
d: [2,250,2,16,7,4,21,249,96]
d: [2,251,2,14,6,5,21,250,106]
d: [2,252,2,12,5,6,21,251,111]
d: [2,253,2,10,4,7,21,252,116]
d: [2,254,2,8,3,8,21,253,121]
d: [2,255,2,6,2,9,21,254,126]
d: [2,256,2,4,1,10,21,255,131]
a: [3,260,2,4,1,256]
s: 260
d: [2,742,2,4,1,2,22,260,152]
d: [2,743,2,0,0,2,21,742,72]
e: 743
s: 342
d: [2,473,0,5,1,2,23,76,172]
a: [3,488,0,5,1,473]
s: 488
a: [3,449,2,4,1,446]
s: 449
d: [2,476,0,5,1,2,23,185,172]
a: [3,498,0,5,1,476]
s: 498
d: [2,480,0,5,1,2,23,308,172]
a: [3,518,0,5,1,480]
s: 518
d: [2,481,0,5,1,2,23,337,172]
a: [3,523,0,5,1,481]
s: 523
d: [2,482,0,5,1,2,23,342,172]
a: [3,528,0,5,1,482]
s: 528
d: [2,753,0,5,1,2,22,488,152]
a: [3,757,0,5,1,753]
s: 757
s: 81
s: 219
s: 179
a: [3,282,0,6,1,279]
s: 282
a: [3,230,1,5,1,227]
s: 230
d: [2,318,0,6,1,2,22,81,152]
a: [3,322,0,6,1,318]
s: 322
r: 152
r: 171
s: 151
s: 96
s: 76
s: 163
s: 185
a: [3,167,1,3,1,148]
s: 167
s: 267
s: 273
s: 172
r: 76
r: 86
r: 185
r: 219
r: 302
r: 308
r: 322
s: 337
s: 449
s: 342
s: 347
d: [2,452,2,28,10,1,40,393,61]
d: [2,453,2,24,9,2,21,452,101]
d: [2,454,2,20,8,3,21,453,86]
d: [2,455,2,16,7,4,21,454,96]
d: [2,456,2,14,6,5,21,455,106]
d: [2,457,2,12,5,6,21,456,111]
d: [2,458,2,10,4,7,21,457,116]
d: [2,459,2,8,3,8,21,458,121]
d: [2,460,2,6,2,9,21,459,126]
d: [2,461,2,4,1,10,21,460,131]
a: [3,465,2,4,1,461]
s: 465
s: 357
d: [2,483,2,4,1,2,23,450,172]
a: [3,532,2,4,1,483]
s: 532
s: 488
s: 498
d: [2,865,2,4,1,2,22,465,172]
a: [3,869,2,4,1,865]
s: 869
s: 523
s: 528
s: 659
s: 666
s: 672
s: 678
s: 684
s: 690
s: 81
s: 197
s: 179
s: 208
s: 282
a: [3,370,0,6,1,367]
s: 370
s: 230
d: [2,475,0,6,1,2,23,179,172]
a: [3,493,0,6,1,475]
s: 493
d: [2,477,1,5,1,2,23,208,172]
a: [3,503,1,5,1,477]
s: 503
d: [2,603,0,6,1,2,22,179,160]
a: [3,607,0,6,1,603]
s: 607
d: [2,770,0,6,1,2,22,179,152]
a: [3,774,0,6,1,770]
s: 774
d: [2,478,1,5,1,2,23,219,172]
a: [3,508,1,5,1,478]
s: 508
d: [2,776,0,6,1,2,22,282,152]
a: [3,780,0,6,1,776]
s: 780
d: [2,479,1,5,1,2,23,231,172]
a: [3,512,1,5,1,479]
s: 512
d: [2,846,0,6,1,2,23,322,172]
a: [3,855,0,6,1,846]
s: 855
d: [2,893,0,6,1,2,22,81,172]
a: [3,897,0,6,1,893]
s: 897
d: [2,962,2,0,0,6,52,106,101,96,91,91,512,897,59]
e: 962
d: [2,963,1,23,6,1,40,897,59]
d: [2,970,1,19,5,1,18,963]
d: [2,972,1,14,4,2,21,970,512]
d: [2,973,1,10,3,3,21,972,91]
d: [2,974,1,6,2,4,21,973,96]
d: [2,975,1,2,1,5,21,974,101]
d: [2,976,1,0,0,6,21,975,106]
e: 976
r: 167
r: 172
r: 897
s: 168
s: 86
s: 171
s: 76
s: 185
s: 449
s: 302
s: 465
d: [2,997,2,4,1,2,22,465,168]
d: [2,998,2,0,0,2,21,997,72]
e: 998
s: 308
s: 337
s: 342
s: 81
s: 208
s: 179
s: 322
s: 219
s: 370
s: 230
s: 774
d: [2,910,0,6,1,2,22,370,172]
a: [3,914,0,6,1,910]
s: 914
d: [2,566,1,5,1,2,23,197,152]
a: [3,574,1,5,1,566]
s: 574
d: [2,922,0,6,1,2,22,774,172]
a: [3,926,0,6,1,922]
s: 926
d: [2,609,1,23,6,1,40,283,59]
d: [2,616,1,19,5,1,18,609]
d: [2,617,1,15,4,2,21,616,101]
d: [2,618,1,11,3,3,21,617,91]
d: [2,619,1,7,2,4,21,618,96]
d: [2,620,1,5,1,5,21,619,106]
a: [3,623,1,5,1,620]
s: 623
d: [2,1009,0,6,1,2,22,81,168]
a: [3,1013,0,6,1,1009]
s: 1013
d: [2,1084,2,4,1,2,22,450,168]
d: [2,1085,2,0,0,2,21,1084,71]
e: 1085
d: [2,1087,2,4,1,2,22,260,168]
d: [2,1088,2,0,0,2,21,1087,71]
e: 1088
r: 163
r: 168
r: 503
r: 512
r: 855
r: 914
s: 86
s: 164
r: 91
r: 273
r: 282
r: 291
r: 488
r: 493
r: 498
r: 508
r: 518
r: 523
r: 528
r: 532
r: 659
r: 672
r: 757
r: 780
r: 869
r: 926
r: 1013
d: [2,1131,2,4,1,2,23,869,164]
d: [2,1146,2,0,0,2,21,1131,72]
e: 1146
s: 76
s: 167
s: 185
s: 302
s: 171
s: 308
s: 260
s: 337
s: 342
s: 450
s: 81
s: 465
s: 179
a: [3,321,0,6,1,318]
s: 321
s: 209
s: 371
s: 219
s: 774
a: [3,854,0,6,1,846]
s: 854
s: 231
d: [2,1017,0,6,1,2,22,179,168]
a: [3,1021,0,6,1,1017]
s: 1021
s: 393
d: [2,1030,0,6,1,2,22,370,168]
a: [3,1033,0,6,1,1030]
s: 1033
d: [2,1038,0,6,1,2,22,774,168]
a: [3,1042,0,6,1,1038]
s: 1042
d: [2,375,1,23,6,1,40,371,57]
d: [2,378,1,19,5,1,18,375]
d: [2,395,1,15,4,2,21,378,86]
d: [2,396,1,11,3,3,21,395,96]
d: [2,397,1,7,2,4,21,396,101]
d: [2,398,1,5,1,5,21,397,106]
a: [3,402,1,5,1,398]
s: 402
d: [2,1244,2,0,0,9,52,131,126,121,116,111,106,96,96,86,72,402,61]
e: 1244
d: [2,1245,2,28,10,1,40,402,61]
d: [2,1246,2,24,9,1,18,1245]
d: [2,1248,2,20,8,1,21,1246,72]
d: [2,1249,2,16,7,2,21,1248,86]
d: [2,1250,2,12,6,3,21,1249,96]
d: [2,1251,2,10,5,4,21,1250,106]
d: [2,1252,2,8,4,5,21,1251,111]
d: [2,1253,2,6,3,6,21,1252,116]
d: [2,1254,2,4,2,7,21,1253,121]
d: [2,1255,2,2,1,8,21,1254,126]
d: [2,1256,2,0,0,9,21,1255,131]
e: 1256
r: 151
r: 164
r: 209
r: 260
r: 371
r: 393
r: 402
r: 1033
s: 91
s: 152
r: 96
r: 167
r: 179
r: 185
r: 197
r: 231
r: 267
r: 337
r: 342
r: 347
r: 357
r: 450
r: 465
r: 607
r: 623
r: 666
r: 678
r: 684
r: 690
r: 1021
r: 1042
d: [2,1283,2,4,1,2,23,450,152]
d: [2,1294,2,0,0,2,21,1283,70]
e: 1294
d: [2,1284,2,4,1,2,23,465,152]
d: [2,1296,2,0,0,2,21,1284,71]
e: 1296
r: 152
s: 96
s: 76
s: 151
s: 185
s: 163
s: 267
s: 273
s: 167
s: 291
s: 172
r: 76
r: 86
r: 185
r: 219
r: 308
r: 774
s: 302
s: 337
s: 450
s: 342
a: [3,464,2,4,1,461]
s: 464
s: 347
s: 357
a: [3,533,2,4,1,483]
s: 533
s: 488
a: [3,868,2,4,1,865]
s: 868
s: 498
s: 518
s: 523
s: 528
s: 659
s: 666
s: 672
s: 678
s: 684
s: 690
s: 757
s: 81
s: 197
s: 179
s: 208
s: 282
s: 322
s: 231
s: 370
a: [3,392,1,5,1,389]
s: 392
s: 493
s: 607
a: [3,401,1,5,1,398]
s: 401
s: 780
s: 503
s: 855
s: 897
s: 508
s: 914
a: [3,513,1,5,1,479]
s: 513
s: 926
a: [3,1012,0,6,1,1009]
s: 1012
s: 623
s: 1021
d: [2,610,1,23,6,1,40,283,58]
d: [2,615,1,19,5,1,18,610]
d: [2,626,1,15,4,2,21,615,101]
d: [2,627,1,11,3,3,21,626,91]
d: [2,628,1,7,2,4,21,627,96]
d: [2,629,1,5,1,5,21,628,106]
a: [3,633,1,5,1,629]
s: 633
d: [2,1469,2,0,0,9,52,131,126,121,116,111,106,91,101,101,70,633,61]
e: 1469
d: [2,1470,2,28,10,1,40,633,61]
d: [2,1471,2,24,9,1,18,1470]
d: [2,1473,2,20,8,1,21,1471,70]
d: [2,1474,2,16,7,2,21,1473,101]
d: [2,1475,2,12,6,3,21,1474,91]
d: [2,1476,2,10,5,4,21,1475,106]
d: [2,1477,2,8,4,5,21,1476,111]
d: [2,1478,2,6,3,6,21,1477,116]
d: [2,1479,2,4,2,7,21,1478,121]
d: [2,1480,2,2,1,8,21,1479,126]
d: [2,1481,2,0,0,9,21,1480,131]
e: 1481
r: 633
r: 855
s: 854
a: [3,1034,0,6,1,1030]
s: 1034
a: [3,632,1,5,1,629]
s: 632
s: 1042
d: [2,611,1,23,6,1,40,283,57]
d: [2,614,1,19,5,1,18,611]
d: [2,635,1,15,4,2,21,614,101]
d: [2,636,1,11,3,3,21,635,91]
d: [2,637,1,7,2,4,21,636,96]
d: [2,638,1,5,1,5,21,637,106]
a: [3,641,1,5,1,638]
s: 641
d: [2,1117,0,6,1,2,23,282,164]
a: [3,1138,0,6,1,1117]
s: 1138
d: [2,1130,0,6,1,2,23,780,164]
a: [3,1144,0,6,1,1130]
s: 1144
d: [2,612,1,23,6,1,40,283,56]
d: [2,613,1,19,5,1,18,612]
d: [2,644,1,15,4,2,21,613,101]
d: [2,645,1,11,3,3,21,644,91]
d: [2,646,1,7,2,4,21,645,96]
d: [2,647,1,5,1,5,21,646,106]
a: [3,651,1,5,1,647]
s: 651
d: [2,1167,0,6,1,2,22,321,164]
a: [3,1171,0,6,1,1167]
s: 1171
d: [2,1506,2,28,10,1,40,651,61]
d: [2,1507,2,24,9,2,21,1506,91]
d: [2,1508,2,20,8,3,21,1507,96]
d: [2,1509,2,16,7,4,21,1508,101]
d: [2,1510,2,14,6,5,21,1509,106]
d: [2,1511,2,12,5,6,21,1510,111]
d: [2,1512,2,10,4,7,21,1511,116]
d: [2,1513,2,8,3,8,21,1512,121]
d: [2,1514,2,6,2,9,21,1513,126]
d: [2,1515,2,4,1,10,21,1514,131]
a: [3,1519,2,4,1,1515]
s: 1519
d: [2,1204,0,6,1,2,22,854,164]
a: [3,1208,0,6,1,1204]
s: 1208
d: [2,1521,0,6,1,2,22,1171,172]
a: [3,1525,0,6,1,1521]
s: 1525
d: [2,704,1,5,1,2,22,208,168]
a: [3,708,1,5,1,704]
s: 708
d: [2,789,1,23,6,1,40,322,58]
d: [2,794,1,19,5,1,18,789]
d: [2,801,1,15,4,2,21,794,86]
d: [2,802,1,11,3,3,21,801,91]
d: [2,803,1,7,2,4,21,802,101]
d: [2,804,1,5,1,5,21,803,106]
a: [3,808,1,5,1,804]
s: 808
d: [2,790,1,23,6,1,40,322,57]
d: [2,793,1,19,5,1,18,790]
d: [2,810,1,15,4,2,21,793,86]
d: [2,811,1,11,3,3,21,810,91]
d: [2,812,1,7,2,4,21,811,101]
d: [2,813,1,5,1,5,21,812,106]
a: [3,817,1,5,1,813]
s: 817
d: [2,791,1,23,6,1,40,322,56]
d: [2,792,1,19,5,1,18,791]
d: [2,819,1,15,4,2,21,792,86]
d: [2,820,1,11,3,3,21,819,91]
d: [2,821,1,7,2,4,21,820,101]
d: [2,822,1,5,1,5,21,821,106]
a: [3,826,1,5,1,822]
s: 826
d: [2,1535,2,0,0,9,52,131,126,121,116,111,106,101,101,91,72,826,61]
e: 1535
d: [2,1536,2,28,10,1,40,826,61]
d: [2,1537,2,24,9,1,18,1536]
d: [2,1539,2,20,8,1,21,1537,72]
d: [2,1540,2,16,7,2,21,1539,91]
d: [2,1541,2,12,6,3,21,1540,101]
d: [2,1542,2,10,5,4,21,1541,106]
d: [2,1543,2,8,4,5,21,1542,111]
d: [2,1544,2,6,3,6,21,1543,116]
d: [2,1545,2,4,2,7,21,1544,121]
d: [2,1546,2,2,1,8,21,1545,126]
d: [2,1547,2,0,0,9,21,1546,131]
e: 1547
r: 826
a: [3,825,1,5,1,822]
s: 825
d: [2,931,1,23,6,1,40,855,58]
d: [2,936,1,15,4,1,18,931]
d: [2,942,1,11,3,2,21,936,91]
d: [2,943,1,7,2,3,21,942,101]
d: [2,944,1,5,1,4,21,943,106]
a: [3,948,1,5,1,944]
s: 948
d: [2,1549,2,0,0,9,52,131,126,121,116,111,106,91,91,101,72,948,61]
e: 1549
d: [2,1550,2,28,10,1,40,948,61]
d: [2,1551,2,24,9,1,18,1550]
d: [2,1553,2,20,8,1,21,1551,72]
d: [2,1554,2,16,7,2,21,1553,101]
d: [2,1555,2,12,6,3,21,1554,91]
d: [2,1556,2,10,5,4,21,1555,106]
d: [2,1557,2,8,4,5,21,1556,111]
d: [2,1558,2,6,3,6,21,1557,116]
d: [2,1559,2,4,2,7,21,1558,121]
d: [2,1560,2,2,1,8,21,1559,126]
d: [2,1561,2,0,0,9,21,1560,131]
e: 1561
d: [2,601,0,6,1,2,22,81,160]
d: [2,768,0,6,1,2,22,81,152]
d: [2,1015,1,5,1,2,22,208,168]
d: [2,1048,1,23,6,1,40,1013,59]
d: [2,1055,1,19,5,1,18,1048]
d: [2,1056,1,15,4,2,21,1055,96]
d: [2,1057,1,11,3,3,21,1056,91]
d: [2,1058,1,7,2,4,21,1057,101]
d: [2,1059,1,5,1,5,21,1058,106]
d: [2,1163,0,6,1,2,22,81,164]
d: [2,1176,1,23,6,1,40,371,57]
d: [2,1179,1,19,5,1,18,1176]
d: [2,1192,1,15,4,2,21,1179,86]
d: [2,1193,1,11,3,3,21,1192,96]
d: [2,1194,1,7,2,4,21,1193,101]
d: [2,1195,1,5,1,5,21,1194,106]
d: [2,1298,1,18,6,1,40,76,65]
d: [2,1299,1,15,5,2,21,1298,159]
d: [2,1300,1,12,4,3,21,1299,155]
d: [2,1361,0,6,1,2,22,81,172]
d: [2,1419,1,23,6,1,40,897,58]
d: [2,1424,1,19,5,1,18,1419]
d: [2,1431,1,15,4,2,21,1424,91]
d: [2,1432,1,11,3,3,21,1431,96]
d: [2,1433,1,7,2,4,21,1432,101]
d: [2,1434,1,5,1,5,21,1433,106]
f: 1561,1547,1481,1434,1361,1300,1256,1195,1163,1088,1059,1015,1009,822,768,629,601,435,256,216,174,43,44,45,46,47,48,49,50,51,52,53,54
% Refutation found. Thanks to Tanya!
% SZS status Theorem for projpl_1__t27_projpl_1
% SZS output start Proof for projpl_1__t27_projpl_1
1. ! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => ! [X4] : (m1_subset_1(X4,u1_incsp_1(X0)) => (r5_projpl_1(X0,X1,X2,X3,X4) => r2_zfmisc_1(X1,X2,X3,X4))))))) [input t27_projpl_1]
2. ~! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => ! [X4] : (m1_subset_1(X4,u1_incsp_1(X0)) => (r5_projpl_1(X0,X1,X2,X3,X4) => r2_zfmisc_1(X1,X2,X3,X4))))))) [negated conjecture 1]
3. ! [X0] : ! [X1] : ! [X2] : (r1_zfmisc_1(X0,X1,X2) <=> (X1 != X2 & X0 != X2 & X0 != X1)) [input d5_zfmisc_1]
4. ! [X0] : (l1_incsp_1(X0) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => ! [X4] : (m1_subset_1(X4,u1_incsp_1(X0)) => (r5_projpl_1(X0,X1,X2,X3,X4) <=> (~r4_projpl_1(X0,X4,X1,X2) & ~r4_projpl_1(X0,X3,X4,X1) & ~r4_projpl_1(X0,X2,X3,X4) & ~r4_projpl_1(X0,X1,X2,X3)))))))) [input d6_projpl_1]
5. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_zfmisc_1(X0,X1,X2,X3) <=> (X2 != X3 & X1 != X3 & X1 != X2 & X0 != X3 & X0 != X2 & X0 != X1)) [input d6_zfmisc_1]
13. ! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (~r4_projpl_1(X0,X1,X2,X3) => r1_zfmisc_1(X1,X2,X3)))))) [input t26_projpl_1]
14. ! [X0,X1,X2] : (r1_zfmisc_1(X0,X1,X2) <=> (X1 != X2 & X0 != X2 & X0 != X1)) [flattening 3]
15. ! [X0,X1,X2,X3] : (r2_zfmisc_1(X0,X1,X2,X3) <=> (X2 != X3 & X1 != X3 & X1 != X2 & X0 != X3 & X0 != X2 & X0 != X1)) [flattening 5]
16. ! [X0,X1,X2] : (r1_zfmisc_1(X0,X1,X2) => (X1 != X2 & X0 != X2 & X0 != X1)) [unused predicate definition removal 14]
17. ! [X0,X1,X2,X3] : ((X2 != X3 & X1 != X3 & X1 != X2 & X0 != X3 & X0 != X2 & X0 != X1) => r2_zfmisc_1(X0,X1,X2,X3)) [unused predicate definition removal 15]
18. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : ((~r2_zfmisc_1(X1,X2,X3,X4) & r5_projpl_1(X0,X1,X2,X3,X4)) & m1_subset_1(X4,u1_incsp_1(X0))) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & (l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0))) [ennf transformation 2]
19. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r2_zfmisc_1(X1,X2,X3,X4) & r5_projpl_1(X0,X1,X2,X3,X4) & m1_subset_1(X4,u1_incsp_1(X0))) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) [flattening 18]
21. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r5_projpl_1(X0,X1,X2,X3,X4) <=> (~r4_projpl_1(X0,X4,X1,X2) & ~r4_projpl_1(X0,X3,X4,X1) & ~r4_projpl_1(X0,X2,X3,X4) & ~r4_projpl_1(X0,X1,X2,X3))) | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [ennf transformation 4]
22. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_zfmisc_1(X1,X2,X3) | r4_projpl_1(X0,X1,X2,X3)) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | (~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0))) [ennf transformation 13]
23. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (r1_zfmisc_1(X1,X2,X3) | r4_projpl_1(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0)) [flattening 22]
24. ! [X0,X1,X2] : ((X1 != X2 & X0 != X2 & X0 != X1) | ~r1_zfmisc_1(X0,X1,X2)) [ennf transformation 16]
25. ! [X0,X1,X2,X3] : (r2_zfmisc_1(X0,X1,X2,X3) | (X2 = X3 | X1 = X3 | X1 = X2 | X0 = X3 | X0 = X2 | X0 = X1)) [ennf transformation 17]
26. ! [X0,X1,X2,X3] : (r2_zfmisc_1(X0,X1,X2,X3) | X2 = X3 | X1 = X3 | X1 = X2 | X0 = X3 | X0 = X2 | X0 = X1) [flattening 25]
27. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r2_zfmisc_1(X1,X2,X3,X4) & r5_projpl_1(X0,X1,X2,X3,X4) & m1_subset_1(X4,u1_incsp_1(X0))) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r2_zfmisc_1(X1,X2,X3,X4) & r5_projpl_1(sK0,X1,X2,X3,X4) & m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v4_incproj(sK0) & v3_incproj(sK0) & v2_incproj(sK0) & v1_incproj(sK0) & v6_incsp_1(sK0)) [choice axiom]
28. ? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r2_zfmisc_1(X1,X2,X3,X4) & r5_projpl_1(sK0,X1,X2,X3,X4) & m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) => (? [X2] : (? [X3] : (? [X4] : (~r2_zfmisc_1(sK1,X2,X3,X4) & r5_projpl_1(sK0,sK1,X2,X3,X4) & m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) [choice axiom]
29. ? [X2] : (? [X3] : (? [X4] : (~r2_zfmisc_1(sK1,X2,X3,X4) & r5_projpl_1(sK0,sK1,X2,X3,X4) & m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) => (? [X3] : (? [X4] : (~r2_zfmisc_1(sK1,sK2,X3,X4) & r5_projpl_1(sK0,sK1,sK2,X3,X4) & m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) [choice axiom]
30. ? [X3] : (? [X4] : (~r2_zfmisc_1(sK1,sK2,X3,X4) & r5_projpl_1(sK0,sK1,sK2,X3,X4) & m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u1_incsp_1(sK0))) => (? [X4] : (~r2_zfmisc_1(sK1,sK2,sK3,X4) & r5_projpl_1(sK0,sK1,sK2,sK3,X4) & m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(sK3,u1_incsp_1(sK0))) [choice axiom]
31. ? [X4] : (~r2_zfmisc_1(sK1,sK2,sK3,X4) & r5_projpl_1(sK0,sK1,sK2,sK3,X4) & m1_subset_1(X4,u1_incsp_1(sK0))) => (~r2_zfmisc_1(sK1,sK2,sK3,sK4) & r5_projpl_1(sK0,sK1,sK2,sK3,sK4) & m1_subset_1(sK4,u1_incsp_1(sK0))) [choice axiom]
32. ((((~r2_zfmisc_1(sK1,sK2,sK3,sK4) & r5_projpl_1(sK0,sK1,sK2,sK3,sK4) & m1_subset_1(sK4,u1_incsp_1(sK0))) & m1_subset_1(sK3,u1_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v4_incproj(sK0) & v3_incproj(sK0) & v2_incproj(sK0) & v1_incproj(sK0) & v6_incsp_1(sK0) [skolemisation 19,31,30,29,28,27]
33. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r5_projpl_1(X0,X1,X2,X3,X4) | (r4_projpl_1(X0,X4,X1,X2) | r4_projpl_1(X0,X3,X4,X1) | r4_projpl_1(X0,X2,X3,X4) | r4_projpl_1(X0,X1,X2,X3))) & ((~r4_projpl_1(X0,X4,X1,X2) & ~r4_projpl_1(X0,X3,X4,X1) & ~r4_projpl_1(X0,X2,X3,X4) & ~r4_projpl_1(X0,X1,X2,X3)) | ~r5_projpl_1(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [nnf transformation 21]
34. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r5_projpl_1(X0,X1,X2,X3,X4) | r4_projpl_1(X0,X4,X1,X2) | r4_projpl_1(X0,X3,X4,X1) | r4_projpl_1(X0,X2,X3,X4) | r4_projpl_1(X0,X1,X2,X3)) & ((~r4_projpl_1(X0,X4,X1,X2) & ~r4_projpl_1(X0,X3,X4,X1) & ~r4_projpl_1(X0,X2,X3,X4) & ~r4_projpl_1(X0,X1,X2,X3)) | ~r5_projpl_1(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [flattening 33]
43. v6_incsp_1(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
44. v1_incproj(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
45. v2_incproj(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
46. v3_incproj(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
47. v4_incproj(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
48. l1_incsp_1(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
49. m1_subset_1(sK1,u1_incsp_1(sK0)) [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
50. m1_subset_1(sK2,u1_incsp_1(sK0)) [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
51. m1_subset_1(sK3,u1_incsp_1(sK0)) [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
52. m1_subset_1(sK4,u1_incsp_1(sK0)) [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
53. r5_projpl_1(sK0,sK1,sK2,sK3,sK4) [cnf transformation 32] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
54. ~r2_zfmisc_1(sK1,sK2,sK3,sK4) [cnf transformation 32] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
56. ~r5_projpl_1(X0,X1,X2,X3,X4) | ~r4_projpl_1(X0,X1,X2,X3) | ~m1_subset_1(X4,u1_incsp_1(X0)) | ~m1_subset_1(X3,u1_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) [cnf transformation 34] {a:0,w:29,wCS:11600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. ~r5_projpl_1(X0,X1,X2,X3,X4) | ~r4_projpl_1(X0,X2,X3,X4) | ~m1_subset_1(X4,u1_incsp_1(X0)) | ~m1_subset_1(X3,u1_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) [cnf transformation 34] {a:0,w:29,wCS:11600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. ~r5_projpl_1(X0,X1,X2,X3,X4) | ~r4_projpl_1(X0,X3,X4,X1) | ~m1_subset_1(X4,u1_incsp_1(X0)) | ~m1_subset_1(X3,u1_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) [cnf transformation 34] {a:0,w:29,wCS:11600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. ~r5_projpl_1(X0,X1,X2,X3,X4) | ~r4_projpl_1(X0,X4,X1,X2) | ~m1_subset_1(X4,u1_incsp_1(X0)) | ~m1_subset_1(X3,u1_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) [cnf transformation 34] {a:0,w:29,wCS:11600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. r4_projpl_1(X0,X1,X2,X3) | r1_zfmisc_1(X1,X2,X3) | ~m1_subset_1(X3,u1_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0) [cnf transformation 23] {a:0,w:33,wCS:13200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. X0 != X1 | ~r1_zfmisc_1(X0,X1,X2) [cnf transformation 24] {a:0,w:7,wCS:2800,sine:2,thAx:0,allAx:1,thDist:-1}
63. X0 != X2 | ~r1_zfmisc_1(X0,X1,X2) [cnf transformation 24] {a:0,w:7,wCS:2800,sine:2,thAx:0,allAx:1,thDist:-1}
64. X1 != X2 | ~r1_zfmisc_1(X0,X1,X2) [cnf transformation 24] {a:0,w:7,wCS:2800,sine:2,thAx:0,allAx:1,thDist:-1}
65. r2_zfmisc_1(X0,X1,X2,X3) | X2 = X3 | X1 = X3 | X1 = X2 | X0 = X3 | X0 = X2 | X0 = X1 [cnf transformation 26] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. ~r1_zfmisc_1(X0,X2,X2) [equality resolution 64] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
71. ~r1_zfmisc_1(X2,X1,X2) [equality resolution 63] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
72. ~r1_zfmisc_1(X1,X1,X2) [equality resolution 62] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
74. 1 <=> r2_zfmisc_1(sK1,sK2,sK3,sK4) [avatar definition]
76. ~r2_zfmisc_1(sK1,sK2,sK3,sK4) <- (~1) [avatar component clause 74] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
77. ~1 [avatar split clause 54,74]
79. 2 <=> r5_projpl_1(sK0,sK1,sK2,sK3,sK4) [avatar definition]
81. r5_projpl_1(sK0,sK1,sK2,sK3,sK4) <- (2) [avatar component clause 79] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
82. 2 [avatar split clause 53,79]
84. 3 <=> m1_subset_1(sK4,u1_incsp_1(sK0)) [avatar definition]
86. m1_subset_1(sK4,u1_incsp_1(sK0)) <- (3) [avatar component clause 84] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
87. 3 [avatar split clause 52,84]
89. 4 <=> m1_subset_1(sK3,u1_incsp_1(sK0)) [avatar definition]
91. m1_subset_1(sK3,u1_incsp_1(sK0)) <- (4) [avatar component clause 89] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
92. 4 [avatar split clause 51,89]
94. 5 <=> m1_subset_1(sK2,u1_incsp_1(sK0)) [avatar definition]
96. m1_subset_1(sK2,u1_incsp_1(sK0)) <- (5) [avatar component clause 94] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
97. 5 [avatar split clause 50,94]
99. 6 <=> m1_subset_1(sK1,u1_incsp_1(sK0)) [avatar definition]
101. m1_subset_1(sK1,u1_incsp_1(sK0)) <- (6) [avatar component clause 99] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
102. 6 [avatar split clause 49,99]
104. 7 <=> l1_incsp_1(sK0) [avatar definition]
106. l1_incsp_1(sK0) <- (7) [avatar component clause 104] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
107. 7 [avatar split clause 48,104]
109. 8 <=> v4_incproj(sK0) [avatar definition]
111. v4_incproj(sK0) <- (8) [avatar component clause 109] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
112. 8 [avatar split clause 47,109]
114. 9 <=> v3_incproj(sK0) [avatar definition]
116. v3_incproj(sK0) <- (9) [avatar component clause 114] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
117. 9 [avatar split clause 46,114]
119. 10 <=> v2_incproj(sK0) [avatar definition]
121. v2_incproj(sK0) <- (10) [avatar component clause 119] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
122. 10 [avatar split clause 45,119]
124. 11 <=> v1_incproj(sK0) [avatar definition]
126. v1_incproj(sK0) <- (11) [avatar component clause 124] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
127. 11 [avatar split clause 44,124]
129. 12 <=> v6_incsp_1(sK0) [avatar definition]
131. v6_incsp_1(sK0) <- (12) [avatar component clause 129] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
132. 12 [avatar split clause 43,129]
150. 16 <=> sK1 = sK2 [avatar definition]
152. sK1 = sK2 <- (16) [avatar component clause 150] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
154. 17 <=> sK1 = sK3 [avatar definition]
155. sK1 != sK3 <- (~17) [avatar component clause 154] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
156. sK1 = sK3 <- (17) [avatar component clause 154] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
158. 18 <=> sK1 = sK4 [avatar definition]
159. sK1 != sK4 <- (~18) [avatar component clause 158] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
160. sK1 = sK4 <- (18) [avatar component clause 158] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
162. 19 <=> sK2 = sK3 [avatar definition]
164. sK2 = sK3 <- (19) [avatar component clause 162] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
166. 20 <=> sK2 = sK4 [avatar definition]
168. sK2 = sK4 <- (20) [avatar component clause 166] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
170. 21 <=> sK3 = sK4 [avatar definition]
172. sK3 = sK4 <- (21) [avatar component clause 170] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
174. r5_projpl_1(sK0,sK1,sK2,sK1,sK4) <- (2, 17) [backward demodulation 81,156] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
178. 22 <=> r5_projpl_1(sK0,sK1,sK2,sK1,sK4) [avatar definition]
180. r5_projpl_1(sK0,sK1,sK2,sK1,sK4) <- (22) [avatar component clause 178] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
181. 22 | ~2 | ~17 [avatar split clause 174,154,79,178]
207. 25 <=> r4_projpl_1(sK0,sK2,sK1,sK4) [avatar definition]
208. r4_projpl_1(sK0,sK2,sK1,sK4) <- (25) [avatar component clause 207] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
209. ~r4_projpl_1(sK0,sK2,sK1,sK4) <- (~25) [avatar component clause 207] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
211. ~r4_projpl_1(sK0,sK1,sK4,sK1) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (22) [resolution 58,180] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
212. ~r4_projpl_1(sK0,sK1,sK4,sK1) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (22) [duplicate literal removal 211] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
213. ~r4_projpl_1(sK0,sK1,sK4,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, 22) [subsumption resolution 212,86] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
214. ~r4_projpl_1(sK0,sK1,sK4,sK1) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, 6, 22) [subsumption resolution 213,101] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
215. ~r4_projpl_1(sK0,sK1,sK4,sK1) | ~l1_incsp_1(sK0) <- (3, 5, 6, 22) [subsumption resolution 214,96] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
216. ~r4_projpl_1(sK0,sK1,sK4,sK1) <- (3, 5, 6, 7, 22) [subsumption resolution 215,106] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
218. 26 <=> r4_projpl_1(sK0,sK1,sK4,sK1) [avatar definition]
220. ~r4_projpl_1(sK0,sK1,sK4,sK1) <- (~26) [avatar component clause 218] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
221. ~26 | ~3 | ~5 | ~6 | ~7 | ~22 [avatar split clause 216,178,104,99,94,84,218]
235. r1_zfmisc_1(sK2,sK1,sK4) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~25) [resolution 61,209] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
248. r1_zfmisc_1(sK2,sK1,sK4) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, ~25) [subsumption resolution 235,86] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
249. r1_zfmisc_1(sK2,sK1,sK4) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 6, ~25) [subsumption resolution 248,101] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
250. r1_zfmisc_1(sK2,sK1,sK4) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, 6, ~25) [subsumption resolution 249,96] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
251. r1_zfmisc_1(sK2,sK1,sK4) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, 6, 7, ~25) [subsumption resolution 250,106] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
252. r1_zfmisc_1(sK2,sK1,sK4) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, 6, 7, 8, ~25) [subsumption resolution 251,111] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
253. r1_zfmisc_1(sK2,sK1,sK4) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, 6, 7, 8, 9, ~25) [subsumption resolution 252,116] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
254. r1_zfmisc_1(sK2,sK1,sK4) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, 6, 7, 8, 9, 10, ~25) [subsumption resolution 253,121] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
255. r1_zfmisc_1(sK2,sK1,sK4) | ~v6_incsp_1(sK0) <- (3, 5, 6, 7, 8, 9, 10, 11, ~25) [subsumption resolution 254,126] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
256. r1_zfmisc_1(sK2,sK1,sK4) <- (3, 5, 6, 7, 8, 9, 10, 11, 12, ~25) [subsumption resolution 255,131] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
258. 28 <=> r1_zfmisc_1(sK2,sK1,sK4) [avatar definition]
260. r1_zfmisc_1(sK2,sK1,sK4) <- (28) [avatar component clause 258] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
261. 28 | ~3 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 25 [avatar split clause 256,207,129,124,119,114,109,104,99,94,84,258]
281. 31 <=> r5_projpl_1(sK0,sK1,sK2,sK3,sK1) [avatar definition]
283. r5_projpl_1(sK0,sK1,sK2,sK3,sK1) <- (31) [avatar component clause 281] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
320. 36 <=> r5_projpl_1(sK0,sK1,sK1,sK3,sK4) [avatar definition]
322. r5_projpl_1(sK0,sK1,sK1,sK3,sK4) <- (36) [avatar component clause 320] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
369. 41 <=> r5_projpl_1(sK0,sK1,sK2,sK2,sK4) [avatar definition]
371. r5_projpl_1(sK0,sK1,sK2,sK2,sK4) <- (41) [avatar component clause 369] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
400. 43 <=> r4_projpl_1(sK0,sK2,sK2,sK4) [avatar definition]
402. ~r4_projpl_1(sK0,sK2,sK2,sK4) <- (~43) [avatar component clause 400] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
424. r1_zfmisc_1(sK1,sK4,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~26) [resolution 220,61] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
425. r1_zfmisc_1(sK1,sK4,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~26) [duplicate literal removal 424] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
427. ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~26) [subsumption resolution 425,71] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
428. ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (6, ~26) [subsumption resolution 427,101] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
429. ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 6, ~26) [subsumption resolution 428,86] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
430. ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 6, 7, ~26) [subsumption resolution 429,106] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
431. ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 6, 7, 8, ~26) [subsumption resolution 430,111] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
432. ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 6, 7, 8, 9, ~26) [subsumption resolution 431,116] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
433. ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 6, 7, 8, 9, 10, ~26) [subsumption resolution 432,121] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
434. ~v6_incsp_1(sK0) <- (3, 6, 7, 8, 9, 10, 11, ~26) [subsumption resolution 433,126] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
435. $false <- (3, 6, 7, 8, 9, 10, 11, 12, ~26) [subsumption resolution 434,131] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
436. ~3 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 26 [avatar contradiction clause 435]
601. r5_projpl_1(sK0,sK1,sK2,sK3,sK1) <- (2, 18) [forward demodulation 81,160] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
602. 31 | ~2 | ~18 [avatar split clause 601,158,79,281]
610. ~r4_projpl_1(sK0,sK3,sK1,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (31) [resolution 283,58] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
615. ~r4_projpl_1(sK0,sK3,sK1,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (31) [duplicate literal removal 610] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
626. ~r4_projpl_1(sK0,sK3,sK1,sK1) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (6, 31) [subsumption resolution 615,101] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
627. ~r4_projpl_1(sK0,sK3,sK1,sK1) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 6, 31) [subsumption resolution 626,91] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
628. ~r4_projpl_1(sK0,sK3,sK1,sK1) | ~l1_incsp_1(sK0) <- (4, 5, 6, 31) [subsumption resolution 627,96] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
629. ~r4_projpl_1(sK0,sK3,sK1,sK1) <- (4, 5, 6, 7, 31) [subsumption resolution 628,106] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
631. 60 <=> r4_projpl_1(sK0,sK3,sK1,sK1) [avatar definition]
633. ~r4_projpl_1(sK0,sK3,sK1,sK1) <- (~60) [avatar component clause 631] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
634. ~60 | ~4 | ~5 | ~6 | ~7 | ~31 [avatar split clause 629,281,104,99,94,89,631]
706. 69 <=> r4_projpl_1(sK0,sK2,sK1,sK2) [avatar definition]
768. r5_projpl_1(sK0,sK1,sK1,sK3,sK4) <- (2, 16) [forward demodulation 81,152] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
769. 36 | ~2 | ~16 [avatar split clause 768,150,79,320]
791. ~r4_projpl_1(sK0,sK1,sK1,sK3) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (36) [resolution 322,56] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
792. ~r4_projpl_1(sK0,sK1,sK1,sK3) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (36) [duplicate literal removal 791] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
819. ~r4_projpl_1(sK0,sK1,sK1,sK3) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, 36) [subsumption resolution 792,86] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
820. ~r4_projpl_1(sK0,sK1,sK1,sK3) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, 4, 36) [subsumption resolution 819,91] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
821. ~r4_projpl_1(sK0,sK1,sK1,sK3) | ~l1_incsp_1(sK0) <- (3, 4, 6, 36) [subsumption resolution 820,101] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
822. ~r4_projpl_1(sK0,sK1,sK1,sK3) <- (3, 4, 6, 7, 36) [subsumption resolution 821,106] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
824. 77 <=> r4_projpl_1(sK0,sK1,sK1,sK3) [avatar definition]
826. ~r4_projpl_1(sK0,sK1,sK1,sK3) <- (~77) [avatar component clause 824] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
827. ~77 | ~3 | ~4 | ~6 | ~7 | ~36 [avatar split clause 822,320,104,99,89,84,824]
895. 80 <=> r5_projpl_1(sK0,sK1,sK2,sK3,sK3) [avatar definition]
897. r5_projpl_1(sK0,sK1,sK2,sK3,sK3) <- (80) [avatar component clause 895] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
946. 83 <=> r4_projpl_1(sK0,sK3,sK3,sK1) [avatar definition]
948. ~r4_projpl_1(sK0,sK3,sK3,sK1) <- (~83) [avatar component clause 946] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1009. r5_projpl_1(sK0,sK1,sK2,sK3,sK2) <- (2, 20) [forward demodulation 81,168] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1011. 86 <=> r5_projpl_1(sK0,sK1,sK2,sK3,sK2) [avatar definition]
1013. r5_projpl_1(sK0,sK1,sK2,sK3,sK2) <- (86) [avatar component clause 1011] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1014. 86 | ~2 | ~20 [avatar split clause 1009,166,79,1011]
1015. r4_projpl_1(sK0,sK2,sK1,sK2) <- (20, 25) [forward demodulation 208,168] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1016. 69 | ~20 | ~25 [avatar split clause 1015,207,166,706]
1048. ~r4_projpl_1(sK0,sK2,sK1,sK2) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (86) [resolution 1013,59] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1055. ~r4_projpl_1(sK0,sK2,sK1,sK2) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (86) [duplicate literal removal 1048] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1056. ~r4_projpl_1(sK0,sK2,sK1,sK2) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (5, 86) [subsumption resolution 1055,96] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1057. ~r4_projpl_1(sK0,sK2,sK1,sK2) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 5, 86) [subsumption resolution 1056,91] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1058. ~r4_projpl_1(sK0,sK2,sK1,sK2) | ~l1_incsp_1(sK0) <- (4, 5, 6, 86) [subsumption resolution 1057,101] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1059. ~r4_projpl_1(sK0,sK2,sK1,sK2) <- (4, 5, 6, 7, 86) [subsumption resolution 1058,106] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1060. ~69 | ~4 | ~5 | ~6 | ~7 | ~86 [avatar split clause 1059,1011,104,99,94,89,706]
1087. r1_zfmisc_1(sK2,sK1,sK2) <- (20, 28) [forward demodulation 260,168] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1088. $false <- (20, 28) [subsumption resolution 1087,71] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
1089. ~20 | ~28 [avatar contradiction clause 1088]
1163. r5_projpl_1(sK0,sK1,sK2,sK2,sK4) <- (2, 19) [forward demodulation 81,164] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1164. 41 | ~2 | ~19 [avatar split clause 1163,162,79,369]
1176. ~r4_projpl_1(sK0,sK2,sK2,sK4) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (41) [resolution 371,57] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1179. ~r4_projpl_1(sK0,sK2,sK2,sK4) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (41) [duplicate literal removal 1176] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1192. ~r4_projpl_1(sK0,sK2,sK2,sK4) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, 41) [subsumption resolution 1179,86] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1193. ~r4_projpl_1(sK0,sK2,sK2,sK4) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, 5, 41) [subsumption resolution 1192,96] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1194. ~r4_projpl_1(sK0,sK2,sK2,sK4) | ~l1_incsp_1(sK0) <- (3, 5, 6, 41) [subsumption resolution 1193,101] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1195. ~r4_projpl_1(sK0,sK2,sK2,sK4) <- (3, 5, 6, 7, 41) [subsumption resolution 1194,106] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1196. ~43 | ~3 | ~5 | ~6 | ~7 | ~41 [avatar split clause 1195,369,104,99,94,84,400]
1245. r1_zfmisc_1(sK2,sK2,sK4) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~43) [resolution 402,61] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1246. r1_zfmisc_1(sK2,sK2,sK4) | ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~43) [duplicate literal removal 1245] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1248. ~m1_subset_1(sK4,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~43) [subsumption resolution 1246,72] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1249. ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, ~43) [subsumption resolution 1248,86] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1250. ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, ~43) [subsumption resolution 1249,96] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1251. ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, 7, ~43) [subsumption resolution 1250,106] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1252. ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, 7, 8, ~43) [subsumption resolution 1251,111] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1253. ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, 7, 8, 9, ~43) [subsumption resolution 1252,116] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1254. ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (3, 5, 7, 8, 9, 10, ~43) [subsumption resolution 1253,121] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1255. ~v6_incsp_1(sK0) <- (3, 5, 7, 8, 9, 10, 11, ~43) [subsumption resolution 1254,126] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1256. $false <- (3, 5, 7, 8, 9, 10, 11, 12, ~43) [subsumption resolution 1255,131] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1257. ~3 | ~5 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 43 [avatar contradiction clause 1256]
1298. sK3 = sK4 | sK2 = sK4 | sK2 = sK3 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 <- (~1) [resolution 76,65] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1299. sK3 = sK4 | sK2 = sK4 | sK2 = sK3 | sK1 = sK3 | sK1 = sK2 <- (~1, ~18) [subsumption resolution 1298,159] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1300. sK3 = sK4 | sK2 = sK4 | sK2 = sK3 | sK1 = sK2 <- (~1, ~17, ~18) [subsumption resolution 1299,155] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1301. 16 | 19 | 20 | 21 | 1 | 17 | 18 [avatar split clause 1300,158,154,74,170,166,162,150]
1361. r5_projpl_1(sK0,sK1,sK2,sK3,sK3) <- (2, 21) [forward demodulation 81,172] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1362. 80 | ~2 | ~21 [avatar split clause 1361,170,79,895]
1419. ~r4_projpl_1(sK0,sK3,sK3,sK1) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (80) [resolution 897,58] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1424. ~r4_projpl_1(sK0,sK3,sK3,sK1) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (80) [duplicate literal removal 1419] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1431. ~r4_projpl_1(sK0,sK3,sK3,sK1) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 80) [subsumption resolution 1424,91] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1432. ~r4_projpl_1(sK0,sK3,sK3,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 5, 80) [subsumption resolution 1431,96] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1433. ~r4_projpl_1(sK0,sK3,sK3,sK1) | ~l1_incsp_1(sK0) <- (4, 5, 6, 80) [subsumption resolution 1432,101] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1434. ~r4_projpl_1(sK0,sK3,sK3,sK1) <- (4, 5, 6, 7, 80) [subsumption resolution 1433,106] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1435. ~83 | ~4 | ~5 | ~6 | ~7 | ~80 [avatar split clause 1434,895,104,99,94,89,946]
1470. r1_zfmisc_1(sK3,sK1,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~60) [resolution 633,61] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1471. r1_zfmisc_1(sK3,sK1,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~60) [duplicate literal removal 1470] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1473. ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~60) [subsumption resolution 1471,70] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1474. ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (6, ~60) [subsumption resolution 1473,101] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1475. ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, ~60) [subsumption resolution 1474,91] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1476. ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, ~60) [subsumption resolution 1475,106] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1477. ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, 8, ~60) [subsumption resolution 1476,111] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1478. ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, 8, 9, ~60) [subsumption resolution 1477,116] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1479. ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, 8, 9, 10, ~60) [subsumption resolution 1478,121] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1480. ~v6_incsp_1(sK0) <- (4, 6, 7, 8, 9, 10, 11, ~60) [subsumption resolution 1479,126] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1481. $false <- (4, 6, 7, 8, 9, 10, 11, 12, ~60) [subsumption resolution 1480,131] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1482. ~4 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 60 [avatar contradiction clause 1481]
1536. r1_zfmisc_1(sK1,sK1,sK3) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~77) [resolution 826,61] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1537. r1_zfmisc_1(sK1,sK1,sK3) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~77) [duplicate literal removal 1536] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1539. ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~77) [subsumption resolution 1537,72] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1540. ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, ~77) [subsumption resolution 1539,91] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1541. ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, ~77) [subsumption resolution 1540,101] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1542. ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, ~77) [subsumption resolution 1541,106] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1543. ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, 8, ~77) [subsumption resolution 1542,111] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1544. ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, 8, 9, ~77) [subsumption resolution 1543,116] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1545. ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, 8, 9, 10, ~77) [subsumption resolution 1544,121] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1546. ~v6_incsp_1(sK0) <- (4, 6, 7, 8, 9, 10, 11, ~77) [subsumption resolution 1545,126] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1547. $false <- (4, 6, 7, 8, 9, 10, 11, 12, ~77) [subsumption resolution 1546,131] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1548. ~4 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 77 [avatar contradiction clause 1547]
1550. r1_zfmisc_1(sK3,sK3,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~83) [resolution 948,61] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1551. r1_zfmisc_1(sK3,sK3,sK1) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~83) [duplicate literal removal 1550] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1553. ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~83) [subsumption resolution 1551,72] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1554. ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (6, ~83) [subsumption resolution 1553,101] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1555. ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, ~83) [subsumption resolution 1554,91] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1556. ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, ~83) [subsumption resolution 1555,106] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1557. ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, 8, ~83) [subsumption resolution 1556,111] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1558. ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, 8, 9, ~83) [subsumption resolution 1557,116] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1559. ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4, 6, 7, 8, 9, 10, ~83) [subsumption resolution 1558,121] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1560. ~v6_incsp_1(sK0) <- (4, 6, 7, 8, 9, 10, 11, ~83) [subsumption resolution 1559,126] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1561. $false <- (4, 6, 7, 8, 9, 10, 11, 12, ~83) [subsumption resolution 1560,131] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1562. ~4 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 83 [avatar contradiction clause 1561]
1563. $false [avatar sat refutation 77,82,87,92,97,102,107,112,117,122,127,132,181,221,261,436,602,634,769,827,1014,1016,1060,1089,1164,1196,1257,1301,1362,1435,1482,1548,1562] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:2628,thDist:-2628}
% SZS output end Proof for projpl_1__t27_projpl_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% >>> Preprocessing
% Introduced skolems: 9
% Unused predicate definitions: 2
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 1236
% Active clauses: 223
% Passive clauses: 342
% Final active clauses: 79
% Final passive clauses: 5
% Main loop iterations started: 312
% 
% >>> Simplifying Inferences
% Duplicate literals: 323
% Fw subsumption resolutions: 325
% Fw demodulations: 98
% Bw demodulations: 102
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Equational tautologies: 67
% Forward subsumptions: 116
% 
% >>> Generating Inferences
% Binary resolution: 163
% Unit resulting resolution: 7
% Backward superposition: 8
% 
% >>> AVATAR
% Split clauses: 28
% Split components: 99
% Sat splitting refutations: 22
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 238
% SAT solver unit clauses: 15
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 5884
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% backward demodulation index maintenance: 0.001 s
% forward demodulation: 0.001 s
% unit resulting resolution: 0.002 s
% other: 0.010 s ( own 0.005 s ) 

