i: [1,85,0,8,2,1,0,0] t17_mcart_1
a: [3,136,0,4,1,85]
s: 136
i: [1,83,0,8,1,1,0,0] t17_mcart_1
a: [3,146,0,8,1,83]
s: 146
i: [1,86,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,151,0,2,1,86]
s: 151
i: [1,87,0,2,1,0,0,255] fc1_xboole_0
a: [3,156,0,2,1,87]
s: 156
i: [1,116,0,2,1,0,0,2] rc2_xboole_0
a: [3,166,0,2,1,116]
s: 166
i: [1,117,0,2,1,0,0,2] rc1_xboole_0
a: [3,171,0,2,1,117]
s: 171
i: [1,119,0,2,1,0,0,2] rc1_relat_1
a: [3,176,0,2,1,119]
s: 176
i: [1,118,0,2,1,0,0,2] rc1_relat_1
a: [3,181,0,2,1,118]
s: 181
i: [1,120,0,2,1,0,0,255] rc2_relat_1
a: [3,186,0,2,1,120]
s: 186
i: [1,89,0,3,1,0,0,1] fc2_xboole_0
s: 89
i: [1,88,0,3,1,0,0,255] d2_xboole_0
a: [3,161,0,3,1,88]
s: 161
r: 151
i: [1,90,0,4,2,0,0,2] cc1_relat_1
s: 90
i: [1,110,0,4,1,0,0,255] existence_m1_subset_1
s: 110
d: [2,189,1,2,1,1,40,90,171]
a: [3,194,1,2,1,189]
s: 194
i: [1,111,0,4,1,0,0,1] fc3_xboole_0
s: 111
i: [1,112,0,4,1,0,0,1] fc6_relat_1
s: 112
d: [2,190,1,2,1,1,40,90,156]
a: [3,199,1,2,1,190]
s: 199
i: [1,122,0,4,1,0,0,1] d1_tarski
s: 122
i: [1,91,0,5,2,0,0,255] t6_boole
s: 91
i: [1,101,0,5,2,0,0,1] t7_boole
s: 101
i: [1,125,0,5,1,0,0,1] d2_tarski
s: 125
d: [2,204,1,7,1,1,40,101,146]
a: [3,209,1,7,1,204]
s: 209
i: [1,127,0,5,1,0,0,1] d2_tarski
s: 127
i: [1,95,0,6,2,0,0,255] t1_subset
s: 95
i: [1,114,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 114
i: [1,100,0,7,3,0,0,1] t8_boole
s: 100
d: [2,213,1,8,1,1,40,95,146]
a: [3,220,1,8,1,213]
s: 220
i: [1,113,0,7,1,0,0,1] commutativity_k2_tarski
s: 113
d: [2,222,1,8,1,1,40,114,146]
a: [3,229,1,8,1,222]
s: 229
i: [1,123,0,7,2,0,0,1] d1_tarski
s: 123
i: [1,94,0,8,3,0,0,255] t2_subset
s: 94
d: [2,214,1,4,1,0,40,95,122]
s: 214
i: [1,115,0,8,3,0,0,1] fc10_subset_1
s: 115
d: [2,223,1,4,1,0,40,114,122]
s: 223
i: [1,102,0,9,2,0,0,1] t10_mcart_1
s: 102
i: [1,103,0,9,2,0,0,1] t10_mcart_1
s: 103
d: [2,247,1,5,1,1,40,103,146]
a: [3,251,1,5,1,247]
s: 251
d: [2,253,2,0,0,2,52,136,251,123]
e: 253
d: [2,254,2,4,1,1,40,251,123]
d: [2,259,2,0,0,2,21,254,136]
e: 259
r: 136
a: [3,132,0,4,1,85]
s: 132
a: [3,135,0,4,1,85]
s: 135
r: 251
i: [1,84,0,8,2,1,0,0] t17_mcart_1
a: [3,141,0,4,1,84]
s: 141
i: [1,128,0,11,3,0,0,1] d2_tarski
s: 128
d: [2,256,2,5,1,1,40,251,95]
a: [3,265,2,5,1,256]
s: 265
i: [1,98,0,14,3,0,0,1] d1_tarski
s: 98
d: [2,255,2,5,1,1,40,251,114]
a: [3,270,2,5,1,255]
s: 270
i: [1,99,0,14,3,0,0,1] d1_tarski
s: 99
i: [1,108,0,17,3,0,0,1] d2_tarski
s: 108
d: [2,241,1,6,1,1,40,102,146]
a: [3,245,1,6,1,241]
s: 245
d: [2,286,2,0,0,3,52,132,141,245,128]
e: 286
d: [2,287,2,8,2,1,40,245,128]
d: [2,292,2,4,1,2,21,287,141]
d: [2,293,2,0,0,3,21,292,132]
e: 293
f: 293,254,247,241,83,84,85
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t17_mcart_1
% SZS output start Proof for mcart_1__t17_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))) => (k2_xtuple_0(X0) = X3 & (k1_xtuple_0(X0) = X2 | k1_xtuple_0(X0) = X1))) [input t17_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))) => (k2_xtuple_0(X0) = X3 & (k1_xtuple_0(X0) = X2 | k1_xtuple_0(X0) = X1))) [negated conjecture 1]
8. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
9. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
29. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [input t10_mcart_1]
35. ~! [X0,X1,X2,X3] : (r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))) => (k2_xtuple_0(X0) = X3 & (k1_xtuple_0(X0) = X2 | k1_xtuple_0(X0) = X1))) [flattening 2]
38. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 8]
41. ! [X0,X1,X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [flattening 29]
42. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 9]
46. ? [X0,X1,X2,X3] : ((k2_xtuple_0(X0) != X3 | (k1_xtuple_0(X0) != X2 & k1_xtuple_0(X0) != X1)) & r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3)))) [ennf transformation 35]
58. ! [X0,X1,X2] : ((r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1)) | ~r2_hidden(X0,k2_zfmisc_1(X1,X2))) [ennf transformation 41]
62. ? [X0,X1,X2,X3] : ((k2_xtuple_0(X0) != X3 | (k1_xtuple_0(X0) != X2 & k1_xtuple_0(X0) != X1)) & r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3)))) => ((sK3 != k2_xtuple_0(sK0) | (sK2 != k1_xtuple_0(sK0) & sK1 != k1_xtuple_0(sK0))) & r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k1_tarski(sK3)))) [choice axiom]
63. (sK3 != k2_xtuple_0(sK0) | (sK2 != k1_xtuple_0(sK0) & sK1 != k1_xtuple_0(sK0))) & r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k1_tarski(sK3))) [skolemisation 46,62]
64. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 38]
65. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 64]
66. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
67. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 65,66]
68. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 42]
69. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 68]
70. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 69]
71. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK5(X0,X1,X2) != X1 & sK5(X0,X1,X2) != X0) | ~r2_hidden(sK5(X0,X1,X2),X2)) & (sK5(X0,X1,X2) = X1 | sK5(X0,X1,X2) = X0 | r2_hidden(sK5(X0,X1,X2),X2)))) [choice axiom]
72. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK5(X0,X1,X2) != X1 & sK5(X0,X1,X2) != X0) | ~r2_hidden(sK5(X0,X1,X2),X2)) & (sK5(X0,X1,X2) = X1 | sK5(X0,X1,X2) = X0 | r2_hidden(sK5(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 70,71]
83. r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k1_tarski(sK3))) [cnf transformation 63] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
84. sK3 != k2_xtuple_0(sK0) | sK1 != k1_xtuple_0(sK0) [cnf transformation 63] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
85. sK3 != k2_xtuple_0(sK0) | sK2 != k1_xtuple_0(sK0) [cnf transformation 63] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
96. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 67] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
102. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k1_xtuple_0(X0),X1) [cnf transformation 58] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
103. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k2_xtuple_0(X0),X2) [cnf transformation 58] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2) | k2_tarski(X0,X1) != X2 [cnf transformation 72] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
123. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 96] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
128. ~r2_hidden(X4,k2_tarski(X0,X1)) | X0 = X4 | X1 = X4 [equality resolution 104] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
130. 1 <=> sK2 = k1_xtuple_0(sK0) [avatar definition]
132. sK2 != k1_xtuple_0(sK0) <- (~1) [avatar component clause 130] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
134. 2 <=> sK3 = k2_xtuple_0(sK0) [avatar definition]
137. ~1 | ~2 [avatar split clause 85,134,130]
139. 3 <=> sK1 = k1_xtuple_0(sK0) [avatar definition]
141. sK1 != k1_xtuple_0(sK0) <- (~3) [avatar component clause 139] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
142. ~3 | ~2 [avatar split clause 84,134,139]
144. 4 <=> r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k1_tarski(sK3))) [avatar definition]
146. r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k1_tarski(sK3))) <- (4) [avatar component clause 144] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
147. 4 [avatar split clause 83,144]
241. r2_hidden(k1_xtuple_0(sK0),k2_tarski(sK1,sK2)) <- (4) [resolution 102,146] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
243. 18 <=> r2_hidden(k1_xtuple_0(sK0),k2_tarski(sK1,sK2)) [avatar definition]
245. r2_hidden(k1_xtuple_0(sK0),k2_tarski(sK1,sK2)) <- (18) [avatar component clause 243] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
246. 18 | ~4 [avatar split clause 241,144,243]
247. r2_hidden(k2_xtuple_0(sK0),k1_tarski(sK3)) <- (4) [resolution 103,146] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
249. 19 <=> r2_hidden(k2_xtuple_0(sK0),k1_tarski(sK3)) [avatar definition]
251. r2_hidden(k2_xtuple_0(sK0),k1_tarski(sK3)) <- (19) [avatar component clause 249] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
252. 19 | ~4 [avatar split clause 247,144,249]
254. sK3 = k2_xtuple_0(sK0) <- (19) [resolution 251,123] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
261. 2 | ~19 [avatar split clause 254,249,134]
287. sK1 = k1_xtuple_0(sK0) | sK2 = k1_xtuple_0(sK0) <- (18) [resolution 245,128] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
292. sK2 = k1_xtuple_0(sK0) <- (~3, 18) [subsumption resolution 287,141] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
293. $false <- (~1, ~3, 18) [subsumption resolution 292,132] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
294. 1 | 3 | ~18 [avatar contradiction clause 293]
295. $false [avatar sat refutation 137,142,147,246,252,261,294] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:111,thDist:-111}
% SZS output end Proof for mcart_1__t17_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 33
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 38
% Generated clauses: 120
% Active clauses: 45
% Passive clauses: 62
% Final active clauses: 42
% Final passive clauses: 15
% Main loop iterations started: 47
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 3
% Fw demodulations: 2
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 3
% Forward subsumptions: 20
% 
% >>> Generating Inferences
% Binary resolution: 37
% Unit resulting resolution: 2
% Forward superposition: 2
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 27
% SAT solver unit clauses: 9
% SAT solver binary clauses: 14
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

