i: [1,35,0,2,1,1,0,0] t89_relat_1
a: [3,62,0,2,1,35]
s: 62
i: [1,36,0,4,1,1,0,0] t89_relat_1
a: [3,57,0,4,1,36]
s: 57
i: [1,37,0,6,1,1,0,0] t89_relat_1
a: [3,52,0,6,1,37]
s: 52
i: [1,44,0,3,1,0,0,1] reflexivity_r1_tarski
s: 44
i: [1,43,0,4,1,0,0,255] existence_m1_subset_1
s: 43
i: [1,45,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 45
i: [1,47,0,6,2,0,0,1] dt_k6_relat_1
s: 47
i: [1,48,0,6,2,0,0,1] fc1_relat_1
s: 48
i: [1,41,0,7,2,0,0,255] t3_subset
s: 41
i: [1,42,0,7,2,0,0,255] t3_subset
s: 42
i: [1,46,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 46
d: [2,65,1,5,1,0,40,41,43]
s: 65
i: [1,38,0,8,3,0,0,255] cc2_relat_1
s: 38
d: [2,67,1,6,2,0,44,48,46]
s: 67
i: [1,40,0,8,2,0,0,1] t28_xboole_1
s: 40
i: [1,39,0,11,2,0,0,1] t88_relat_1
s: 39
d: [2,68,1,6,2,0,44,48,46]
s: 68
d: [2,69,1,6,2,0,40,38,43]
s: 69
d: [2,70,1,7,3,0,40,38,42]
s: 70
d: [2,74,1,8,2,0,44,40,46]
s: 74
d: [2,75,1,8,2,0,44,40,46]
s: 75
d: [2,90,2,5,2,1,40,70,57]
a: [3,100,2,3,1,90]
s: 100
d: [2,77,1,8,2,0,44,46,40]
s: 77
d: [2,76,2,7,3,0,44,67,40]
s: 76
d: [2,78,1,8,2,0,44,46,40]
s: 78
d: [2,80,1,10,2,0,44,39,45]
s: 80
d: [2,86,2,9,3,0,44,67,39]
s: 86
d: [2,84,1,10,2,0,44,45,39]
s: 84
d: [2,102,2,9,3,0,44,74,40]
s: 102
d: [2,89,1,10,3,0,44,48,39]
s: 89
d: [2,82,1,11,2,0,44,39,46]
s: 82
d: [2,118,3,8,2,1,40,102,57]
a: [3,124,3,4,1,118]
s: 124
d: [2,83,1,11,2,0,44,39,46]
s: 83
d: [2,106,2,9,3,0,44,40,74]
s: 106
d: [2,87,1,11,2,0,44,46,39]
s: 87
d: [2,88,1,11,2,0,44,46,39]
s: 88
d: [2,120,3,10,2,0,40,102,65]
s: 120
d: [2,81,1,13,3,0,44,39,40]
s: 81
d: [2,105,2,12,3,0,44,74,39]
s: 105
d: [2,159,3,0,0,3,52,62,57,52,105]
e: 159
d: [2,161,3,9,3,1,44,52,105]
d: [2,171,3,6,2,1,19,161]
d: [2,174,3,2,1,2,21,171,57]
d: [2,175,3,0,0,3,21,174,62]
e: 175
f: 175,35,36,37
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t89_relat_1
% SZS output start Proof for relat_1__t89_relat_1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,k10_xtuple_0(X1)) => k10_xtuple_0(k6_relat_1(X0,X1)) = X0)) [input t89_relat_1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,k10_xtuple_0(X1)) => k10_xtuple_0(k6_relat_1(X0,X1)) = X0)) [negated conjecture 1]
4. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
14. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
16. ! [X0] : ! [X1] : (v1_relat_1(X1) => k10_xtuple_0(k6_relat_1(X0,X1)) = k3_xboole_0(k10_xtuple_0(X1),X0)) [input t88_relat_1]
17. ~! [X0,X1] : (v1_relat_1(X1) => (r1_tarski(X0,k10_xtuple_0(X1)) => k10_xtuple_0(k6_relat_1(X0,X1)) = X0)) [flattening 2]
18. ! [X0,X1] : (v1_relat_1(X1) => k10_xtuple_0(k6_relat_1(X0,X1)) = k3_xboole_0(k10_xtuple_0(X1),X0)) [flattening 16]
19. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 14]
23. ? [X0,X1] : ((k10_xtuple_0(k6_relat_1(X0,X1)) != X0 & r1_tarski(X0,k10_xtuple_0(X1))) & v1_relat_1(X1)) [ennf transformation 17]
24. ? [X0,X1] : (k10_xtuple_0(k6_relat_1(X0,X1)) != X0 & r1_tarski(X0,k10_xtuple_0(X1)) & v1_relat_1(X1)) [flattening 23]
26. ! [X0,X1] : (k10_xtuple_0(k6_relat_1(X0,X1)) = k3_xboole_0(k10_xtuple_0(X1),X0) | ~v1_relat_1(X1)) [ennf transformation 18]
27. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 19]
30. ? [X0,X1] : (k10_xtuple_0(k6_relat_1(X0,X1)) != X0 & r1_tarski(X0,k10_xtuple_0(X1)) & v1_relat_1(X1)) => (sK0 != k10_xtuple_0(k6_relat_1(sK0,sK1)) & r1_tarski(sK0,k10_xtuple_0(sK1)) & v1_relat_1(sK1)) [choice axiom]
31. sK0 != k10_xtuple_0(k6_relat_1(sK0,sK1)) & r1_tarski(sK0,k10_xtuple_0(sK1)) & v1_relat_1(sK1) [skolemisation 24,30]
35. v1_relat_1(sK1) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. r1_tarski(sK0,k10_xtuple_0(sK1)) [cnf transformation 31] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. sK0 != k10_xtuple_0(k6_relat_1(sK0,sK1)) [cnf transformation 31] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. k10_xtuple_0(k6_relat_1(X0,X1)) = k3_xboole_0(k10_xtuple_0(X1),X0) | ~v1_relat_1(X1) [cnf transformation 26] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
40. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 27] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
46. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. 1 <=> sK0 = k10_xtuple_0(k6_relat_1(sK0,sK1)) [avatar definition]
52. sK0 != k10_xtuple_0(k6_relat_1(sK0,sK1)) <- (~1) [avatar component clause 50] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
53. ~1 [avatar split clause 37,50]
55. 2 <=> r1_tarski(sK0,k10_xtuple_0(sK1)) [avatar definition]
57. r1_tarski(sK0,k10_xtuple_0(sK1)) <- (2) [avatar component clause 55] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
58. 2 [avatar split clause 36,55]
60. 3 <=> v1_relat_1(sK1) [avatar definition]
62. v1_relat_1(sK1) <- (3) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. 3 [avatar split clause 35,60]
74. k3_xboole_0(X4,X3) = X3 | ~r1_tarski(X3,X4) [superposition 40,46] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
105. k10_xtuple_0(k6_relat_1(X10,X9)) = X10 | ~r1_tarski(X10,k10_xtuple_0(X9)) | ~v1_relat_1(X9) [superposition 74,39] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
161. sK0 != sK0 | ~r1_tarski(sK0,k10_xtuple_0(sK1)) | ~v1_relat_1(sK1) <- (~1) [superposition 52,105] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
171. ~r1_tarski(sK0,k10_xtuple_0(sK1)) | ~v1_relat_1(sK1) <- (~1) [trivial inequality removal 161] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
174. ~v1_relat_1(sK1) <- (~1, 2) [subsumption resolution 171,57] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
175. $false <- (~1, 2, 3) [subsumption resolution 174,62] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
176. 1 | ~2 | ~3 [avatar contradiction clause 175]
180. $false [avatar sat refutation 53,58,63,176] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:38,thDist:-38}
% SZS output end Proof for relat_1__t89_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 14
% Generated clauses: 115
% Active clauses: 29
% Passive clauses: 53
% Final active clauses: 29
% Final passive clauses: 14
% Main loop iterations started: 39
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Trivial inequalities: 1
% Fw subsumption resolutions: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 1
% Forward subsumptions: 31
% 
% >>> Generating Inferences
% Binary resolution: 10
% Unit resulting resolution: 1
% Forward superposition: 24
% Backward superposition: 40
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

