i: [1,36,0,2,1,1,0,0] t98_relat_1
a: [3,64,0,2,1,36]
s: 64
i: [1,37,0,3,1,1,0,0] t98_relat_1
a: [3,59,0,3,1,37]
s: 59
i: [1,38,0,9,1,1,0,0] t98_relat_1
a: [3,54,0,9,1,38]
s: 54
i: [1,45,0,3,1,0,0,1] reflexivity_r1_tarski
s: 45
i: [1,44,0,4,1,0,0,255] existence_m1_subset_1
s: 44
i: [1,46,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 46
i: [1,48,0,6,2,0,0,2] fc1_relat_1
s: 48
i: [1,49,0,6,2,0,0,1] dt_k6_relat_1
s: 49
i: [1,41,0,7,2,0,0,255] t3_subset
s: 41
i: [1,42,0,7,2,0,0,255] t3_subset
s: 42
i: [1,47,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 47
d: [2,67,1,5,1,0,40,41,44]
s: 67
i: [1,39,0,8,3,0,0,255] cc2_relat_1
s: 39
d: [2,69,1,6,2,0,44,48,47]
s: 69
i: [1,40,0,8,2,0,0,1] t28_xboole_1
s: 40
i: [1,50,0,11,2,0,0,1] rd7_relat_1
s: 50
d: [2,70,1,6,2,0,44,48,47]
s: 70
i: [1,43,0,13,2,0,0,1] t96_relat_1
s: 43
d: [2,71,1,6,2,0,40,39,44]
s: 71
d: [2,72,1,7,3,0,40,39,42]
s: 72
d: [2,76,1,8,2,0,44,40,47]
s: 76
d: [2,96,2,4,2,1,40,72,59]
a: [3,106,2,2,1,96]
s: 106
d: [2,77,1,8,2,0,44,40,47]
s: 77
d: [2,78,2,7,3,0,44,69,40]
s: 78
d: [2,79,1,8,2,0,44,47,40]
s: 79
d: [2,80,1,8,2,0,44,47,40]
s: 80
d: [2,90,1,10,3,0,44,49,43]
d: [2,92,1,8,2,0,18,90]
s: 92
d: [2,86,1,13,2,0,44,43,47]
s: 86
d: [2,108,2,9,3,0,44,76,40]
s: 108
d: [2,87,1,13,2,0,44,43,47]
s: 87
d: [2,85,1,14,3,0,44,43,40]
s: 85
d: [2,143,3,6,2,1,40,108,59]
a: [3,149,3,3,1,143]
s: 149
d: [2,89,1,19,3,0,44,50,43]
d: [2,93,1,17,2,0,18,89]
s: 93
d: [2,111,2,9,3,0,44,40,76]
s: 111
d: [2,88,1,23,3,0,44,43,50]
d: [2,94,1,17,2,0,21,88,49]
s: 94
d: [2,91,1,23,3,0,44,50,43]
d: [2,95,1,17,2,0,21,91,49]
s: 95
d: [2,118,2,12,3,0,44,92,43]
d: [2,122,2,10,2,0,18,118]
s: 122
d: [2,131,2,12,3,0,44,92,86]
d: [2,138,2,10,2,0,18,131]
s: 138
d: [2,145,3,10,2,0,40,108,67]
s: 145
d: [2,130,2,15,3,0,44,43,86]
d: [2,139,2,13,2,0,18,130]
s: 139
d: [2,128,2,15,3,0,44,86,43]
d: [2,140,2,13,2,0,18,128]
s: 140
d: [2,339,3,11,2,1,44,54,139]
d: [2,400,3,9,1,2,21,339,64]
a: [3,404,3,9,1,400]
s: 404
d: [2,417,4,0,0,3,52,64,59,404,85]
e: 417
d: [2,418,4,12,3,1,44,404,85]
d: [2,419,4,5,2,1,19,418]
d: [2,421,4,3,1,2,21,419,64]
d: [2,422,4,0,0,3,21,421,59]
e: 422
d: [2,356,3,11,2,1,44,54,139]
d: [2,415,3,9,1,2,21,356,64]
f: 422,415,36,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t98_relat_1
% SZS output start Proof for relat_1__t98_relat_1
1. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => k6_relat_1(X0,X2) = k6_relat_1(X1,k6_relat_1(X0,X2)))) [input t98_relat_1]
2. ~! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => k6_relat_1(X0,X2) = k6_relat_1(X1,k6_relat_1(X0,X2)))) [negated conjecture 1]
4. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
14. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
16. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(k3_xboole_0(X0,X1),X2)) [input t96_relat_1]
17. ~! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => k6_relat_1(X0,X2) = k6_relat_1(X1,k6_relat_1(X0,X2)))) [flattening 2]
18. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 14]
20. ! [X0,X1,X2] : (v1_relat_1(X2) => k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(k3_xboole_0(X0,X1),X2)) [flattening 16]
23. ? [X0,X1,X2] : ((k6_relat_1(X0,X2) != k6_relat_1(X1,k6_relat_1(X0,X2)) & r1_tarski(X0,X1)) & v1_relat_1(X2)) [ennf transformation 17]
24. ? [X0,X1,X2] : (k6_relat_1(X0,X2) != k6_relat_1(X1,k6_relat_1(X0,X2)) & r1_tarski(X0,X1) & v1_relat_1(X2)) [flattening 23]
26. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 18]
27. ! [X0,X1,X2] : (k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(k3_xboole_0(X0,X1),X2) | ~v1_relat_1(X2)) [ennf transformation 20]
31. ? [X0,X1,X2] : (k6_relat_1(X0,X2) != k6_relat_1(X1,k6_relat_1(X0,X2)) & r1_tarski(X0,X1) & v1_relat_1(X2)) => (k6_relat_1(sK0,sK2) != k6_relat_1(sK1,k6_relat_1(sK0,sK2)) & r1_tarski(sK0,sK1) & v1_relat_1(sK2)) [choice axiom]
32. k6_relat_1(sK0,sK2) != k6_relat_1(sK1,k6_relat_1(sK0,sK2)) & r1_tarski(sK0,sK1) & v1_relat_1(sK2) [skolemisation 24,31]
36. v1_relat_1(sK2) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. r1_tarski(sK0,sK1) [cnf transformation 32] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. k6_relat_1(sK0,sK2) != k6_relat_1(sK1,k6_relat_1(sK0,sK2)) [cnf transformation 32] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(k3_xboole_0(X0,X1),X2) | ~v1_relat_1(X2) [cnf transformation 27] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
52. 1 <=> k6_relat_1(sK0,sK2) = k6_relat_1(sK1,k6_relat_1(sK0,sK2)) [avatar definition]
54. k6_relat_1(sK0,sK2) != k6_relat_1(sK1,k6_relat_1(sK0,sK2)) <- (~1) [avatar component clause 52] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
55. ~1 [avatar split clause 38,52]
57. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
59. r1_tarski(sK0,sK1) <- (2) [avatar component clause 57] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 2 [avatar split clause 37,57]
62. 3 <=> v1_relat_1(sK2) [avatar definition]
64. v1_relat_1(sK2) <- (3) [avatar component clause 62] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. 3 [avatar split clause 36,62]
85. k6_relat_1(X2,k6_relat_1(X3,X4)) = k6_relat_1(X2,X4) | ~v1_relat_1(X4) | ~r1_tarski(X2,X3) [superposition 43,40] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
86. k6_relat_1(X5,k6_relat_1(X6,X7)) = k6_relat_1(k3_xboole_0(X6,X5),X7) | ~v1_relat_1(X7) [superposition 43,47] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
130. k6_relat_1(X1,k6_relat_1(X0,X2)) = k6_relat_1(X0,k6_relat_1(X1,X2)) | ~v1_relat_1(X2) | ~v1_relat_1(X2) [superposition 43,86] {a:2,w:15,wCS:6000,sine:1,thAx:0,allAx:3,thDist:-3}
139. k6_relat_1(X1,k6_relat_1(X0,X2)) = k6_relat_1(X0,k6_relat_1(X1,X2)) | ~v1_relat_1(X2) [duplicate literal removal 130] {a:2,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
356. k6_relat_1(sK0,sK2) != k6_relat_1(sK0,k6_relat_1(sK1,sK2)) | ~v1_relat_1(sK2) <- (~1) [superposition 54,139] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
402. 9 <=> k6_relat_1(sK0,sK2) = k6_relat_1(sK0,k6_relat_1(sK1,sK2)) [avatar definition]
404. k6_relat_1(sK0,sK2) != k6_relat_1(sK0,k6_relat_1(sK1,sK2)) <- (~9) [avatar component clause 402] {a:3,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
415. k6_relat_1(sK0,sK2) != k6_relat_1(sK0,k6_relat_1(sK1,sK2)) <- (~1, 3) [subsumption resolution 356,64] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
416. ~9 | 1 | ~3 [avatar split clause 415,62,52,402]
418. k6_relat_1(sK0,sK2) != k6_relat_1(sK0,sK2) | ~v1_relat_1(sK2) | ~r1_tarski(sK0,sK1) <- (~9) [superposition 404,85] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
419. ~v1_relat_1(sK2) | ~r1_tarski(sK0,sK1) <- (~9) [trivial inequality removal 418] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
421. ~r1_tarski(sK0,sK1) <- (3, ~9) [subsumption resolution 419,64] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
422. $false <- (2, 3, ~9) [subsumption resolution 421,59] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
423. ~2 | ~3 | 9 [avatar contradiction clause 422]
424. $false [avatar sat refutation 55,60,65,416,423] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:66,thDist:-66}
% SZS output end Proof for relat_1__t98_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 348
% Active clauses: 32
% Passive clauses: 144
% Final active clauses: 32
% Final passive clauses: 102
% Main loop iterations started: 42
% 
% >>> Simplifying Inferences
% Duplicate literals: 39
% Trivial inequalities: 1
% Fw subsumption resolutions: 78
% Fw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Equational tautologies: 1
% Forward subsumptions: 72
% 
% >>> Generating Inferences
% Binary resolution: 10
% Unit resulting resolution: 1
% Forward superposition: 103
% Backward superposition: 81
% Self superposition: 10
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.002 s
% superposition: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

