i: [1,130,0,2,1,1,0,0] t56_tops_1
a: [3,196,0,2,1,130]
s: 196
i: [1,134,0,3,1,1,0,0] t56_tops_1
a: [3,176,0,3,1,134]
s: 176
i: [1,133,0,3,1,1,0,0] t56_tops_1
a: [3,181,0,3,1,133]
s: 181
i: [1,132,0,3,1,1,0,0] t56_tops_1
a: [3,186,0,3,1,132]
s: 186
i: [1,131,0,5,1,1,0,0] t56_tops_1
a: [3,191,0,5,1,131]
s: 191
i: [1,135,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,201,0,2,1,135]
s: 201
i: [1,136,0,2,1,0,0,255] fc1_xboole_0
a: [3,206,0,2,1,136]
s: 206
i: [1,170,0,2,1,0,0,2] rc2_xboole_0
a: [3,216,0,2,1,170]
s: 216
i: [1,171,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,221,0,2,1,171]
s: 221
i: [1,172,0,2,1,0,0,2] rc1_xboole_0
a: [3,226,0,2,1,172]
s: 226
i: [1,138,0,3,1,0,0,2] fc1_subset_1
s: 138
i: [1,163,0,3,1,0,0,2] rc2_subset_1
s: 163
i: [1,166,0,3,1,0,0,255] reflexivity_r1_tarski
s: 166
i: [1,137,0,3,1,0,0,255] d2_xboole_0
a: [3,211,0,3,1,137]
s: 211
r: 201
i: [1,161,0,4,1,0,0,1] existence_m1_subset_1
s: 161
i: [1,165,0,4,1,0,0,255] rc3_subset_1
s: 165
i: [1,142,0,5,2,0,0,2] rc1_subset_1
s: 142
i: [1,151,0,5,2,0,0,255] t6_boole
s: 151
i: [1,158,0,5,2,0,0,255] t7_boole
s: 158
i: [1,162,0,5,1,0,0,2] rc2_subset_1
s: 162
i: [1,164,0,5,1,0,0,255] rc3_subset_1
s: 164
i: [1,144,0,6,2,0,0,255] rc4_subset_1
s: 144
i: [1,150,0,6,2,0,0,1] rc5_tops_1
s: 150
i: [1,154,0,6,2,0,0,255] t1_subset
s: 154
i: [1,167,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 167
i: [1,141,0,7,2,0,0,2] rc1_subset_1
s: 141
i: [1,143,0,7,2,0,0,255] rc4_subset_1
s: 143
i: [1,156,0,7,2,0,0,255] t3_subset
s: 156
i: [1,157,0,7,3,0,0,3] t8_boole
s: 157
i: [1,149,0,8,2,0,0,1] rc5_tops_1
s: 149
i: [1,152,0,8,3,0,0,2] cc1_subset_1
s: 152
d: [2,232,1,5,2,1,40,152,191]
a: [3,242,1,3,1,232]
s: 242
i: [1,155,0,8,3,0,0,255] t2_subset
s: 155
d: [2,249,1,9,2,1,40,155,191]
d: [2,257,1,5,1,1,21,249,138]
a: [3,261,1,5,1,257]
s: 261
i: [1,153,0,9,3,0,0,255] cc4_subset_1
s: 153
i: [1,160,0,9,3,0,0,255] t5_subset
s: 160
d: [2,268,2,5,1,1,40,261,167]
a: [3,274,2,5,1,268]
s: 274
i: [1,159,0,10,3,0,0,255] t4_subset
s: 159
d: [2,293,1,7,2,1,40,159,191]
s: 293
i: [1,139,0,11,4,0,0,255] cc3_subset_1
s: 139
i: [1,140,0,11,4,0,0,255] cc2_subset_1
s: 140
a: [3,245,1,2,1,232]
s: 245
i: [1,146,0,12,4,0,0,1] cc2_tops_1
s: 146
d: [2,300,2,10,3,1,40,293,155]
d: [2,301,2,7,2,2,21,300,242]
s: 301
i: [1,168,0,14,3,0,0,2] dt_k2_pre_topc
s: 168
i: [1,145,0,15,4,0,0,2] t22_pre_topc
s: 145
d: [2,344,1,10,3,1,40,145,191]
d: [2,353,1,7,2,2,21,344,181]
d: [2,354,1,5,1,3,21,353,196]
a: [3,358,1,5,1,354]
s: 358
i: [1,147,0,15,4,0,0,1] d5_tops_1
s: 147
d: [2,330,3,7,2,2,40,301,167]
s: 330
i: [1,148,0,15,4,0,0,1] d5_tops_1
s: 148
d: [2,379,2,13,4,1,44,148,358]
d: [2,380,2,10,3,2,21,379,186]
d: [2,381,2,7,2,3,21,380,176]
d: [2,382,2,2,1,4,21,381,191]
d: [2,383,2,0,0,5,21,382,196]
e: 383
f: 383,354,130,131,132,133,134
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tops_1__t56_tops_1
% SZS output start Proof for tops_1__t56_tops_1
1. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ((v4_pre_topc(X1,X0) & v2_tops_1(X1,X0)) => v3_tops_1(X1,X0)))) [input t56_tops_1]
2. ~! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ((v4_pre_topc(X1,X0) & v2_tops_1(X1,X0)) => v3_tops_1(X1,X0)))) [negated conjecture 1]
14. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v3_tops_1(X1,X0) <=> v2_tops_1(k2_pre_topc(X0,X1),X0)))) [input d5_tops_1]
44. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (((k2_pre_topc(X0,X1) = X1 & v2_pre_topc(X0)) => v4_pre_topc(X1,X0)) & (v4_pre_topc(X1,X0) => k2_pre_topc(X0,X1) = X1)))) [input t22_pre_topc]
74. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v4_pre_topc(X1,X0) => k2_pre_topc(X0,X1) = X1))) [pure predicate removal 44]
77. ? [X0] : (? [X1] : ((~v3_tops_1(X1,X0) & (v4_pre_topc(X1,X0) & v2_tops_1(X1,X0))) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0)) [ennf transformation 2]
78. ? [X0] : (? [X1] : (~v3_tops_1(X1,X0) & v4_pre_topc(X1,X0) & v2_tops_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0)) [flattening 77]
85. ! [X0] : (! [X1] : ((k2_pre_topc(X0,X1) = X1 | ~v4_pre_topc(X1,X0)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 74]
86. ! [X0] : (! [X1] : (k2_pre_topc(X0,X1) = X1 | ~v4_pre_topc(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [flattening 85]
89. ! [X0] : (! [X1] : ((v3_tops_1(X1,X0) <=> v2_tops_1(k2_pre_topc(X0,X1),X0)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 14]
108. ? [X0] : (? [X1] : (~v3_tops_1(X1,X0) & v4_pre_topc(X1,X0) & v2_tops_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0)) => (? [X1] : (~v3_tops_1(X1,sK0) & v4_pre_topc(X1,sK0) & v2_tops_1(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0)) [choice axiom]
109. ? [X1] : (~v3_tops_1(X1,sK0) & v4_pre_topc(X1,sK0) & v2_tops_1(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => (~v3_tops_1(sK1,sK0) & v4_pre_topc(sK1,sK0) & v2_tops_1(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
110. (~v3_tops_1(sK1,sK0) & v4_pre_topc(sK1,sK0) & v2_tops_1(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) [skolemisation 78,109,108]
115. ! [X0] : (! [X1] : (((v3_tops_1(X1,X0) | ~v2_tops_1(k2_pre_topc(X0,X1),X0)) & (v2_tops_1(k2_pre_topc(X0,X1),X0) | ~v3_tops_1(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [nnf transformation 89]
130. l1_pre_topc(sK0) [cnf transformation 110] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
131. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 110] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
132. v2_tops_1(sK1,sK0) [cnf transformation 110] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
133. v4_pre_topc(sK1,sK0) [cnf transformation 110] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
134. ~v3_tops_1(sK1,sK0) [cnf transformation 110] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
145. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v4_pre_topc(X1,X0) | k2_pre_topc(X0,X1) = X1 | ~l1_pre_topc(X0) [cnf transformation 86] {a:0,w:15,wCS:6000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
148. ~v2_tops_1(k2_pre_topc(X0,X1),X0) | v3_tops_1(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 115] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
174. 1 <=> v3_tops_1(sK1,sK0) [avatar definition]
176. ~v3_tops_1(sK1,sK0) <- (~1) [avatar component clause 174] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
177. ~1 [avatar split clause 134,174]
179. 2 <=> v4_pre_topc(sK1,sK0) [avatar definition]
181. v4_pre_topc(sK1,sK0) <- (2) [avatar component clause 179] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
182. 2 [avatar split clause 133,179]
184. 3 <=> v2_tops_1(sK1,sK0) [avatar definition]
186. v2_tops_1(sK1,sK0) <- (3) [avatar component clause 184] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
187. 3 [avatar split clause 132,184]
189. 4 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
191. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (4) [avatar component clause 189] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
192. 4 [avatar split clause 131,189]
194. 5 <=> l1_pre_topc(sK0) [avatar definition]
196. l1_pre_topc(sK0) <- (5) [avatar component clause 194] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
197. 5 [avatar split clause 130,194]
344. ~v4_pre_topc(sK1,sK0) | sK1 = k2_pre_topc(sK0,sK1) | ~l1_pre_topc(sK0) <- (4) [resolution 145,191] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
353. sK1 = k2_pre_topc(sK0,sK1) | ~l1_pre_topc(sK0) <- (2, 4) [subsumption resolution 344,181] {a:1,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
354. sK1 = k2_pre_topc(sK0,sK1) <- (2, 4, 5) [subsumption resolution 353,196] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
356. 17 <=> sK1 = k2_pre_topc(sK0,sK1) [avatar definition]
358. sK1 = k2_pre_topc(sK0,sK1) <- (17) [avatar component clause 356] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
359. 17 | ~2 | ~4 | ~5 [avatar split clause 354,194,189,179,356]
379. ~v2_tops_1(sK1,sK0) | v3_tops_1(sK1,sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (17) [superposition 148,358] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
380. v3_tops_1(sK1,sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (3, 17) [subsumption resolution 379,186] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
381. ~m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (~1, 3, 17) [subsumption resolution 380,176] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
382. ~l1_pre_topc(sK0) <- (~1, 3, 4, 17) [subsumption resolution 381,191] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
383. $false <- (~1, 3, 4, 5, 17) [subsumption resolution 382,196] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
384. 1 | ~3 | ~4 | ~5 | ~17 [avatar contradiction clause 383]
385. $false [avatar sat refutation 177,182,187,192,197,359,384] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:90,thDist:-90}
% SZS output end Proof for tops_1__t56_tops_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 50
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 3
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 43
% Generated clauses: 182
% Active clauses: 49
% Passive clauses: 113
% Final active clauses: 48
% Final passive clauses: 63
% Main loop iterations started: 50
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 19
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Forward subsumptions: 16
% 
% >>> Generating Inferences
% Binary resolution: 89
% Forward superposition: 1
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 11
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% backward subsumption index maintenance: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

