i: [1,45,0,2,1,1,0,0] t101_pboole
a: [3,89,0,2,1,45]
s: 89
i: [1,43,0,2,1,1,0,0] t101_pboole
a: [3,99,0,2,1,43]
s: 99
i: [1,41,0,2,1,1,0,0] t101_pboole
a: [3,109,0,2,1,41]
s: 109
i: [1,39,0,2,1,1,0,0] t101_pboole
a: [3,119,0,2,1,39]
s: 119
i: [1,37,0,2,1,1,0,0] t101_pboole
a: [3,129,0,2,1,37]
s: 129
i: [1,35,0,2,1,1,0,0] t101_pboole
a: [3,139,0,2,1,35]
s: 139
i: [1,33,0,2,1,1,0,0] t101_pboole
a: [3,149,0,2,1,33]
s: 149
i: [1,31,0,2,1,1,0,0] t101_pboole
a: [3,159,0,2,1,31]
s: 159
i: [1,46,0,3,1,1,0,0] t101_pboole
a: [3,84,0,3,1,46]
s: 84
i: [1,44,0,3,1,1,0,0] t101_pboole
a: [3,94,0,3,1,44]
s: 94
i: [1,42,0,3,1,1,0,0] t101_pboole
a: [3,104,0,3,1,42]
s: 104
i: [1,40,0,3,1,1,0,0] t101_pboole
a: [3,114,0,3,1,40]
s: 114
i: [1,38,0,3,1,1,0,0] t101_pboole
a: [3,124,0,3,1,38]
s: 124
i: [1,36,0,3,1,1,0,0] t101_pboole
a: [3,134,0,3,1,36]
s: 134
i: [1,34,0,3,1,1,0,0] t101_pboole
a: [3,144,0,3,1,34]
s: 144
i: [1,32,0,3,1,1,0,0] t101_pboole
a: [3,154,0,3,1,32]
s: 154
i: [1,50,0,4,1,1,0,0] t101_pboole
a: [3,64,0,4,1,50]
s: 64
i: [1,49,0,4,1,1,0,0] t101_pboole
a: [3,69,0,4,1,49]
s: 69
i: [1,48,0,4,1,1,0,0] t101_pboole
a: [3,74,0,4,1,48]
s: 74
i: [1,47,0,4,1,1,0,0] t101_pboole
a: [3,79,0,4,1,47]
s: 79
i: [1,60,0,2,1,0,0,1] rc1_funct_1
a: [3,164,0,2,1,60]
s: 164
i: [1,59,0,2,1,0,0,1] rc1_funct_1
a: [3,169,0,2,1,59]
s: 169
i: [1,53,0,3,1,0,0,1] rc1_pboole
s: 53
i: [1,55,0,3,1,0,0,1] rc1_pboole
s: 55
i: [1,54,0,4,1,0,0,1] rc1_pboole
s: 54
i: [1,56,0,4,1,0,0,1] rc1_pboole
s: 56
i: [1,57,0,24,9,0,0,1] reflexivity_r2_pboole
s: 57
i: [1,58,0,28,10,0,0,1] symmetry_r4_pboole
s: 58
d: [2,171,1,24,9,1,40,58,69]
d: [2,172,1,21,8,2,21,171,104]
d: [2,173,1,19,7,3,21,172,109]
d: [2,174,1,16,6,4,21,173,114]
d: [2,175,1,14,5,5,21,174,119]
d: [2,176,1,11,4,6,21,175,144]
d: [2,177,1,9,3,7,21,176,149]
d: [2,178,1,6,2,8,21,177,154]
d: [2,179,1,4,1,9,21,178,159]
a: [3,183,1,4,1,179]
s: 183
i: [1,51,0,42,15,0,0,1] t99_pboole
s: 51
i: [1,52,0,42,15,0,0,1] t100_pboole
s: 52
d: [2,186,1,38,14,1,40,51,69]
d: [2,188,1,35,13,2,21,186,104]
d: [2,189,1,33,12,3,21,188,109]
d: [2,190,1,30,11,4,21,189,114]
d: [2,191,1,28,10,5,21,190,119]
d: [2,192,1,25,9,6,21,191,144]
d: [2,193,1,23,8,7,21,192,149]
d: [2,194,1,20,7,8,21,193,154]
d: [2,195,1,18,6,9,21,194,159]
s: 195
d: [2,204,1,38,14,1,40,52,69]
d: [2,206,1,35,13,2,21,204,104]
d: [2,207,1,33,12,3,21,206,109]
d: [2,208,1,30,11,4,21,207,114]
d: [2,209,1,28,10,5,21,208,119]
d: [2,210,1,25,9,6,21,209,144]
d: [2,211,1,23,8,7,21,210,149]
d: [2,212,1,20,7,8,21,211,154]
d: [2,213,1,18,6,9,21,212,159]
s: 213
d: [2,222,2,14,5,10,40,195,74]
d: [2,225,2,11,4,11,21,222,84]
d: [2,226,2,9,3,12,21,225,89]
d: [2,227,2,6,2,13,21,226,94]
d: [2,228,2,4,1,14,21,227,99]
a: [3,232,2,4,1,228]
s: 232
d: [2,246,3,0,0,15,52,159,89,149,129,99,139,154,144,94,84,134,124,64,79,232,52]
e: 246
f: 246,228,31,32,33,34,35,36,37,38,39,40,41,42,43,44,45,46,47,48,49,50
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pboole__t101_pboole
% SZS output start Proof for pboole__t101_pboole
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ! [X4] : ((v1_partfun1(X4,X0) & v1_funct_1(X4) & v4_relat_1(X4,X0) & v1_relat_1(X4)) => ((r4_pboole(X0,X1,X3) & r2_pboole(X0,X3,X4) & r2_pboole(X0,X1,X2)) => r4_pboole(X0,X2,X4)))))) [input t101_pboole]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ! [X4] : ((v1_partfun1(X4,X0) & v1_funct_1(X4) & v4_relat_1(X4,X0) & v1_relat_1(X4)) => ((r4_pboole(X0,X1,X3) & r2_pboole(X0,X3,X4) & r2_pboole(X0,X1,X2)) => r4_pboole(X0,X2,X4)))))) [negated conjecture 1]
7. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X1,X3) & r4_pboole(X0,X1,X2)) => r4_pboole(X0,X3,X2))))) [input t100_pboole]
8. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r4_pboole(X0,X1,X2)) => r4_pboole(X0,X1,X3))))) [input t99_pboole]
9. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ! [X4] : ((v1_partfun1(X4,X0) & v1_funct_1(X4) & v4_relat_1(X4,X0) & v1_relat_1(X4)) => ((r4_pboole(X0,X1,X3) & r2_pboole(X0,X3,X4) & r2_pboole(X0,X1,X2)) => r4_pboole(X0,X2,X4)))))) [flattening 2]
10. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r4_pboole(X0,X1,X2)) => r4_pboole(X0,X1,X3))))) [flattening 8]
11. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X1,X3) & r4_pboole(X0,X1,X2)) => r4_pboole(X0,X3,X2))))) [flattening 7]
12. ? [X0,X1] : (? [X2] : (? [X3] : (? [X4] : ((~r4_pboole(X0,X2,X4) & (r4_pboole(X0,X1,X3) & r2_pboole(X0,X3,X4) & r2_pboole(X0,X1,X2))) & (v1_partfun1(X4,X0) & v1_funct_1(X4) & v4_relat_1(X4,X0) & v1_relat_1(X4))) & (v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3))) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 9]
13. ? [X0,X1] : (? [X2] : (? [X3] : (? [X4] : (~r4_pboole(X0,X2,X4) & r4_pboole(X0,X1,X3) & r2_pboole(X0,X3,X4) & r2_pboole(X0,X1,X2) & v1_partfun1(X4,X0) & v1_funct_1(X4) & v4_relat_1(X4,X0) & v1_relat_1(X4)) & v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 12]
14. ! [X0,X1] : (! [X2] : (! [X3] : ((r4_pboole(X0,X1,X3) | (~r2_pboole(X0,X2,X3) | ~r4_pboole(X0,X1,X2))) | (~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 10]
15. ! [X0,X1] : (! [X2] : (! [X3] : (r4_pboole(X0,X1,X3) | ~r2_pboole(X0,X2,X3) | ~r4_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 14]
16. ! [X0,X1] : (! [X2] : (! [X3] : ((r4_pboole(X0,X3,X2) | (~r2_pboole(X0,X1,X3) | ~r4_pboole(X0,X1,X2))) | (~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 11]
17. ! [X0,X1] : (! [X2] : (! [X3] : (r4_pboole(X0,X3,X2) | ~r2_pboole(X0,X1,X3) | ~r4_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 16]
22. ? [X0,X1] : (? [X2] : (? [X3] : (? [X4] : (~r4_pboole(X0,X2,X4) & r4_pboole(X0,X1,X3) & r2_pboole(X0,X3,X4) & r2_pboole(X0,X1,X2) & v1_partfun1(X4,X0) & v1_funct_1(X4) & v4_relat_1(X4,X0) & v1_relat_1(X4)) & v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (? [X4] : (~r4_pboole(sK0,X2,X4) & r4_pboole(sK0,sK1,X3) & r2_pboole(sK0,X3,X4) & r2_pboole(sK0,sK1,X2) & v1_partfun1(X4,sK0) & v1_funct_1(X4) & v4_relat_1(X4,sK0) & v1_relat_1(X4)) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
23. ? [X2] : (? [X3] : (? [X4] : (~r4_pboole(sK0,X2,X4) & r4_pboole(sK0,sK1,X3) & r2_pboole(sK0,X3,X4) & r2_pboole(sK0,sK1,X2) & v1_partfun1(X4,sK0) & v1_funct_1(X4) & v4_relat_1(X4,sK0) & v1_relat_1(X4)) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (? [X3] : (? [X4] : (~r4_pboole(sK0,sK2,X4) & r4_pboole(sK0,sK1,X3) & r2_pboole(sK0,X3,X4) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(X4,sK0) & v1_funct_1(X4) & v4_relat_1(X4,sK0) & v1_relat_1(X4)) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
24. ? [X3] : (? [X4] : (~r4_pboole(sK0,sK2,X4) & r4_pboole(sK0,sK1,X3) & r2_pboole(sK0,X3,X4) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(X4,sK0) & v1_funct_1(X4) & v4_relat_1(X4,sK0) & v1_relat_1(X4)) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) => (? [X4] : (~r4_pboole(sK0,sK2,X4) & r4_pboole(sK0,sK1,sK3) & r2_pboole(sK0,sK3,X4) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(X4,sK0) & v1_funct_1(X4) & v4_relat_1(X4,sK0) & v1_relat_1(X4)) & v1_partfun1(sK3,sK0) & v1_funct_1(sK3) & v4_relat_1(sK3,sK0) & v1_relat_1(sK3)) [choice axiom]
25. ? [X4] : (~r4_pboole(sK0,sK2,X4) & r4_pboole(sK0,sK1,sK3) & r2_pboole(sK0,sK3,X4) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(X4,sK0) & v1_funct_1(X4) & v4_relat_1(X4,sK0) & v1_relat_1(X4)) => (~r4_pboole(sK0,sK2,sK4) & r4_pboole(sK0,sK1,sK3) & r2_pboole(sK0,sK3,sK4) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(sK4,sK0) & v1_funct_1(sK4) & v4_relat_1(sK4,sK0) & v1_relat_1(sK4)) [choice axiom]
26. (((~r4_pboole(sK0,sK2,sK4) & r4_pboole(sK0,sK1,sK3) & r2_pboole(sK0,sK3,sK4) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(sK4,sK0) & v1_funct_1(sK4) & v4_relat_1(sK4,sK0) & v1_relat_1(sK4)) & v1_partfun1(sK3,sK0) & v1_funct_1(sK3) & v4_relat_1(sK3,sK0) & v1_relat_1(sK3)) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 13,25,24,23,22]
31. v1_relat_1(sK1) [cnf transformation 26] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
32. v4_relat_1(sK1,sK0) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
33. v1_funct_1(sK1) [cnf transformation 26] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
34. v1_partfun1(sK1,sK0) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
35. v1_relat_1(sK2) [cnf transformation 26] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
36. v4_relat_1(sK2,sK0) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
37. v1_funct_1(sK2) [cnf transformation 26] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
38. v1_partfun1(sK2,sK0) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
39. v1_relat_1(sK3) [cnf transformation 26] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
40. v4_relat_1(sK3,sK0) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
41. v1_funct_1(sK3) [cnf transformation 26] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
42. v1_partfun1(sK3,sK0) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
43. v1_relat_1(sK4) [cnf transformation 26] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
44. v4_relat_1(sK4,sK0) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
45. v1_funct_1(sK4) [cnf transformation 26] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
46. v1_partfun1(sK4,sK0) [cnf transformation 26] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
47. r2_pboole(sK0,sK1,sK2) [cnf transformation 26] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
48. r2_pboole(sK0,sK3,sK4) [cnf transformation 26] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
49. r4_pboole(sK0,sK1,sK3) [cnf transformation 26] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
50. ~r4_pboole(sK0,sK2,sK4) [cnf transformation 26] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
51. ~r4_pboole(X0,X1,X2) | ~r2_pboole(X0,X2,X3) | r4_pboole(X0,X1,X3) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 15] {a:0,w:42,wCS:16800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. ~r4_pboole(X0,X1,X2) | ~r2_pboole(X0,X1,X3) | r4_pboole(X0,X3,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 17] {a:0,w:42,wCS:16800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. 1 <=> r4_pboole(sK0,sK2,sK4) [avatar definition]
64. ~r4_pboole(sK0,sK2,sK4) <- (~1) [avatar component clause 62] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
65. ~1 [avatar split clause 50,62]
67. 2 <=> r4_pboole(sK0,sK1,sK3) [avatar definition]
69. r4_pboole(sK0,sK1,sK3) <- (2) [avatar component clause 67] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
70. 2 [avatar split clause 49,67]
72. 3 <=> r2_pboole(sK0,sK3,sK4) [avatar definition]
74. r2_pboole(sK0,sK3,sK4) <- (3) [avatar component clause 72] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
75. 3 [avatar split clause 48,72]
77. 4 <=> r2_pboole(sK0,sK1,sK2) [avatar definition]
79. r2_pboole(sK0,sK1,sK2) <- (4) [avatar component clause 77] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
80. 4 [avatar split clause 47,77]
82. 5 <=> v1_partfun1(sK4,sK0) [avatar definition]
84. v1_partfun1(sK4,sK0) <- (5) [avatar component clause 82] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
85. 5 [avatar split clause 46,82]
87. 6 <=> v1_funct_1(sK4) [avatar definition]
89. v1_funct_1(sK4) <- (6) [avatar component clause 87] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
90. 6 [avatar split clause 45,87]
92. 7 <=> v4_relat_1(sK4,sK0) [avatar definition]
94. v4_relat_1(sK4,sK0) <- (7) [avatar component clause 92] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
95. 7 [avatar split clause 44,92]
97. 8 <=> v1_relat_1(sK4) [avatar definition]
99. v1_relat_1(sK4) <- (8) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
100. 8 [avatar split clause 43,97]
102. 9 <=> v1_partfun1(sK3,sK0) [avatar definition]
104. v1_partfun1(sK3,sK0) <- (9) [avatar component clause 102] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
105. 9 [avatar split clause 42,102]
107. 10 <=> v1_funct_1(sK3) [avatar definition]
109. v1_funct_1(sK3) <- (10) [avatar component clause 107] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
110. 10 [avatar split clause 41,107]
112. 11 <=> v4_relat_1(sK3,sK0) [avatar definition]
114. v4_relat_1(sK3,sK0) <- (11) [avatar component clause 112] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
115. 11 [avatar split clause 40,112]
117. 12 <=> v1_relat_1(sK3) [avatar definition]
119. v1_relat_1(sK3) <- (12) [avatar component clause 117] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
120. 12 [avatar split clause 39,117]
122. 13 <=> v1_partfun1(sK2,sK0) [avatar definition]
124. v1_partfun1(sK2,sK0) <- (13) [avatar component clause 122] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
125. 13 [avatar split clause 38,122]
127. 14 <=> v1_funct_1(sK2) [avatar definition]
129. v1_funct_1(sK2) <- (14) [avatar component clause 127] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
130. 14 [avatar split clause 37,127]
132. 15 <=> v4_relat_1(sK2,sK0) [avatar definition]
134. v4_relat_1(sK2,sK0) <- (15) [avatar component clause 132] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
135. 15 [avatar split clause 36,132]
137. 16 <=> v1_relat_1(sK2) [avatar definition]
139. v1_relat_1(sK2) <- (16) [avatar component clause 137] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
140. 16 [avatar split clause 35,137]
142. 17 <=> v1_partfun1(sK1,sK0) [avatar definition]
144. v1_partfun1(sK1,sK0) <- (17) [avatar component clause 142] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
145. 17 [avatar split clause 34,142]
147. 18 <=> v1_funct_1(sK1) [avatar definition]
149. v1_funct_1(sK1) <- (18) [avatar component clause 147] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
150. 18 [avatar split clause 33,147]
152. 19 <=> v4_relat_1(sK1,sK0) [avatar definition]
154. v4_relat_1(sK1,sK0) <- (19) [avatar component clause 152] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
155. 19 [avatar split clause 32,152]
157. 20 <=> v1_relat_1(sK1) [avatar definition]
159. v1_relat_1(sK1) <- (20) [avatar component clause 157] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
160. 20 [avatar split clause 31,157]
186. ~r2_pboole(sK0,sK3,X0) | r4_pboole(sK0,sK1,X0) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2) [resolution 51,69] {a:1,w:38,wCS:3800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
188. ~r2_pboole(sK0,sK3,X0) | r4_pboole(sK0,sK1,X0) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 9) [subsumption resolution 186,104] {a:1,w:35,wCS:3500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
189. ~r2_pboole(sK0,sK3,X0) | r4_pboole(sK0,sK1,X0) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 9, 10) [subsumption resolution 188,109] {a:1,w:33,wCS:3300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
190. ~r2_pboole(sK0,sK3,X0) | r4_pboole(sK0,sK1,X0) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | ~v1_relat_1(sK3) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 9, 10, 11) [subsumption resolution 189,114] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
191. ~r2_pboole(sK0,sK3,X0) | r4_pboole(sK0,sK1,X0) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 9, 10, 11, 12) [subsumption resolution 190,119] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
192. ~r2_pboole(sK0,sK3,X0) | r4_pboole(sK0,sK1,X0) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 9, 10, 11, 12, 17) [subsumption resolution 191,144] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
193. ~r2_pboole(sK0,sK3,X0) | r4_pboole(sK0,sK1,X0) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 9, 10, 11, 12, 17, 18) [subsumption resolution 192,149] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
194. ~r2_pboole(sK0,sK3,X0) | r4_pboole(sK0,sK1,X0) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | ~v1_relat_1(sK1) <- (2, 9, 10, 11, 12, 17, 18, 19) [subsumption resolution 193,154] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
195. ~r2_pboole(sK0,sK3,X0) | r4_pboole(sK0,sK1,X0) | ~v1_partfun1(X0,sK0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (2, 9, 10, 11, 12, 17, 18, 19, 20) [subsumption resolution 194,159] {a:1,w:18,wCS:1800,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
222. r4_pboole(sK0,sK1,sK4) | ~v1_partfun1(sK4,sK0) | ~v1_funct_1(sK4) | ~v4_relat_1(sK4,sK0) | ~v1_relat_1(sK4) <- (2, 3, 9, 10, 11, 12, 17, 18, 19, 20) [resolution 195,74] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
225. r4_pboole(sK0,sK1,sK4) | ~v1_funct_1(sK4) | ~v4_relat_1(sK4,sK0) | ~v1_relat_1(sK4) <- (2, 3, 5, 9, 10, 11, 12, 17, 18, 19, 20) [subsumption resolution 222,84] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
226. r4_pboole(sK0,sK1,sK4) | ~v4_relat_1(sK4,sK0) | ~v1_relat_1(sK4) <- (2, 3, 5, 6, 9, 10, 11, 12, 17, 18, 19, 20) [subsumption resolution 225,89] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
227. r4_pboole(sK0,sK1,sK4) | ~v1_relat_1(sK4) <- (2, 3, 5, 6, 7, 9, 10, 11, 12, 17, 18, 19, 20) [subsumption resolution 226,94] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
228. r4_pboole(sK0,sK1,sK4) <- (2, 3, 5, 6, 7, 8, 9, 10, 11, 12, 17, 18, 19, 20) [subsumption resolution 227,99] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
230. 24 <=> r4_pboole(sK0,sK1,sK4) [avatar definition]
232. r4_pboole(sK0,sK1,sK4) <- (24) [avatar component clause 230] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
233. 24 | ~2 | ~3 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~17 | ~18 | ~19 | ~20 [avatar split clause 228,157,152,147,142,117,112,107,102,97,92,87,82,72,67,230]
246. $false <- (~1, 4, 5, 6, 7, 8, 13, 14, 15, 16, 17, 18, 19, 20, 24) [unit resulting resolution 159,89,149,129,99,139,154,144,94,84,134,124,64,79,232,52] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:82,thDist:-82}
250. 1 | ~4 | ~5 | ~6 | ~7 | ~8 | ~13 | ~14 | ~15 | ~16 | ~17 | ~18 | ~19 | ~20 | ~24 [avatar contradiction clause 246]
275. $false [avatar sat refutation 65,70,75,80,85,90,95,100,105,110,115,120,125,130,135,140,145,150,155,160,233,250] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:276,thDist:-276}
% SZS output end Proof for pboole__t101_pboole
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 7
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 30
% Generated clauses: 143
% Active clauses: 34
% Passive clauses: 37
% Final active clauses: 34
% Final passive clauses: 3
% Main loop iterations started: 34
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Fw subsumption resolutions: 72
% 
% >>> Deletion Inferences
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 28
% SAT solver unit clauses: 22
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

