i: [1,22,0,2,1,1,0,0] t20_wellord1
a: [3,39,0,2,1,22]
s: 39
i: [1,23,0,11,1,1,0,0] t20_wellord1
a: [3,34,0,11,1,23]
s: 34
i: [1,26,0,4,1,0,0,255] fc6_relat_1
s: 26
i: [1,27,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 27
i: [1,29,0,6,2,0,0,2] fc1_relat_1
s: 29
i: [1,30,0,6,2,0,0,1] dt_k2_wellord1
s: 30
i: [1,28,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 28
i: [1,24,0,11,2,0,0,255] d6_wellord1
s: 24
d: [2,42,1,6,2,0,44,29,28]
s: 42
i: [1,25,0,13,2,0,0,1] t19_wellord1
s: 25
d: [2,59,1,13,2,1,44,34,25]
d: [2,62,1,11,1,2,21,59,39]
d: [2,63,1,11,1,2,22,62,28]
a: [3,67,1,11,1,63]
s: 67
d: [2,69,2,0,0,2,52,39,67,25]
e: 69
d: [2,70,2,13,2,1,44,67,25]
d: [2,71,2,2,1,1,19,70]
d: [2,73,2,0,0,2,21,71,39]
e: 73
f: 73,63,22,23
% Refutation found. Thanks to Tanya!
% SZS status Theorem for wellord1__t20_wellord1
% SZS output start Proof for wellord1__t20_wellord1
1. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => k2_wellord1(k2_wellord1(X2,X0),X1) = k2_wellord1(k2_wellord1(X2,X1),X0)) [input t20_wellord1]
2. ~! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => k2_wellord1(k2_wellord1(X2,X0),X1) = k2_wellord1(k2_wellord1(X2,X1),X0)) [negated conjecture 1]
3. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
11. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => k2_wellord1(k2_wellord1(X2,X0),X1) = k2_wellord1(X2,k3_xboole_0(X0,X1))) [input t19_wellord1]
12. ~! [X0,X1,X2] : (v1_relat_1(X2) => k2_wellord1(k2_wellord1(X2,X0),X1) = k2_wellord1(k2_wellord1(X2,X1),X0)) [flattening 2]
13. ! [X0,X1,X2] : (v1_relat_1(X2) => k2_wellord1(k2_wellord1(X2,X0),X1) = k2_wellord1(X2,k3_xboole_0(X0,X1))) [flattening 11]
15. ? [X0,X1,X2] : (k2_wellord1(k2_wellord1(X2,X0),X1) != k2_wellord1(k2_wellord1(X2,X1),X0) & v1_relat_1(X2)) [ennf transformation 12]
17. ! [X0,X1,X2] : (k2_wellord1(k2_wellord1(X2,X0),X1) = k2_wellord1(X2,k3_xboole_0(X0,X1)) | ~v1_relat_1(X2)) [ennf transformation 13]
20. ? [X0,X1,X2] : (k2_wellord1(k2_wellord1(X2,X0),X1) != k2_wellord1(k2_wellord1(X2,X1),X0) & v1_relat_1(X2)) => (k2_wellord1(k2_wellord1(sK2,sK0),sK1) != k2_wellord1(k2_wellord1(sK2,sK1),sK0) & v1_relat_1(sK2)) [choice axiom]
21. k2_wellord1(k2_wellord1(sK2,sK0),sK1) != k2_wellord1(k2_wellord1(sK2,sK1),sK0) & v1_relat_1(sK2) [skolemisation 15,20]
22. v1_relat_1(sK2) [cnf transformation 21] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
23. k2_wellord1(k2_wellord1(sK2,sK0),sK1) != k2_wellord1(k2_wellord1(sK2,sK1),sK0) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
25. k2_wellord1(k2_wellord1(X2,X0),X1) = k2_wellord1(X2,k3_xboole_0(X0,X1)) | ~v1_relat_1(X2) [cnf transformation 17] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
28. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
32. 1 <=> k2_wellord1(k2_wellord1(sK2,sK0),sK1) = k2_wellord1(k2_wellord1(sK2,sK1),sK0) [avatar definition]
34. k2_wellord1(k2_wellord1(sK2,sK0),sK1) != k2_wellord1(k2_wellord1(sK2,sK1),sK0) <- (~1) [avatar component clause 32] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
35. ~1 [avatar split clause 23,32]
37. 2 <=> v1_relat_1(sK2) [avatar definition]
39. v1_relat_1(sK2) <- (2) [avatar component clause 37] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
40. 2 [avatar split clause 22,37]
59. k2_wellord1(k2_wellord1(sK2,sK0),sK1) != k2_wellord1(sK2,k3_xboole_0(sK1,sK0)) | ~v1_relat_1(sK2) <- (~1) [superposition 34,25] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
62. k2_wellord1(k2_wellord1(sK2,sK0),sK1) != k2_wellord1(sK2,k3_xboole_0(sK1,sK0)) <- (~1, 2) [subsumption resolution 59,39] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
63. k2_wellord1(k2_wellord1(sK2,sK0),sK1) != k2_wellord1(sK2,k3_xboole_0(sK0,sK1)) <- (~1, 2) [forward demodulation 62,28] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
65. 3 <=> k2_wellord1(k2_wellord1(sK2,sK0),sK1) = k2_wellord1(sK2,k3_xboole_0(sK0,sK1)) [avatar definition]
67. k2_wellord1(k2_wellord1(sK2,sK0),sK1) != k2_wellord1(sK2,k3_xboole_0(sK0,sK1)) <- (~3) [avatar component clause 65] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
68. ~3 | 1 | ~2 [avatar split clause 63,37,32,65]
70. k2_wellord1(sK2,k3_xboole_0(sK0,sK1)) != k2_wellord1(sK2,k3_xboole_0(sK0,sK1)) | ~v1_relat_1(sK2) <- (~3) [superposition 67,25] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
71. ~v1_relat_1(sK2) <- (~3) [trivial inequality removal 70] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
73. $false <- (2, ~3) [subsumption resolution 71,39] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
74. ~2 | 3 [avatar contradiction clause 73]
75. $false [avatar sat refutation 35,40,68,74] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:22,thDist:-22}
% SZS output end Proof for wellord1__t20_wellord1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 9
% Generated clauses: 39
% Active clauses: 11
% Passive clauses: 19
% Final active clauses: 11
% Final passive clauses: 8
% Main loop iterations started: 11
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Trivial inequalities: 1
% Fw subsumption resolutions: 5
% Fw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Unit resulting resolution: 1
% Forward superposition: 9
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 7
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% superposition: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

