i: [1,32,0,14,5,1,0,0] t23_gate_1
a: [3,62,0,2,1,32]
s: 62
a: [3,45,0,6,1,32]
s: 45
i: [1,33,0,2,1,0,0,255] fc1_xboole_0
a: [3,71,0,2,1,33]
s: 71
i: [1,41,0,2,1,0,0,1] rc2_xboole_0
a: [3,76,0,2,1,41]
s: 76
i: [1,42,0,2,1,0,0,1] rc1_xboole_0
a: [3,81,0,2,1,42]
s: 81
i: [1,34,0,5,2,0,0,255] fc2_gate_1
s: 34
i: [1,35,0,5,2,0,0,255] fc1_gate_1
s: 35
i: [1,37,0,9,2,0,0,255] d17_gate_1
s: 37
a: [3,49,0,2,1,32]
s: 49
i: [1,38,0,9,2,0,0,255] d17_gate_1
s: 38
a: [3,53,0,2,1,32]
s: 53
i: [1,39,0,9,2,0,0,255] d17_gate_1
s: 39
a: [3,57,0,2,1,32]
s: 57
i: [1,40,0,9,2,0,0,255] d17_gate_1
s: 40
d: [2,93,1,4,2,1,44,45,40]
d: [2,94,1,2,1,2,21,93,71]
d: [2,95,1,0,0,3,21,94,62]
e: 95
r: 45
r: 62
a: [3,61,0,2,1,32]
s: 61
a: [3,46,0,6,1,32]
s: 46
i: [1,36,0,16,5,0,0,255] d17_gate_1
s: 36
d: [2,109,1,11,5,1,44,46,36]
d: [2,110,1,9,4,2,21,109,61]
d: [2,111,1,7,3,3,21,110,57]
d: [2,112,1,5,2,4,21,111,53]
d: [2,113,1,3,1,5,21,112,49]
a: [3,117,1,3,1,113]
s: 117
d: [2,119,2,0,0,2,52,71,117,35]
e: 119
d: [2,120,2,2,1,1,40,117,35]
d: [2,122,2,0,0,2,21,120,71]
e: 122
i: [1,31,0,8,2,1,0,0] t23_gate_1
i: [1,30,0,8,2,1,0,0] t23_gate_1
i: [1,29,0,8,2,1,0,0] t23_gate_1
i: [1,28,0,8,2,1,0,0] t23_gate_1
d: [2,84,1,4,2,1,44,45,37]
d: [2,85,1,2,1,2,21,84,71]
d: [2,87,1,4,2,1,44,45,38]
d: [2,88,1,2,1,2,21,87,71]
d: [2,90,1,4,2,1,44,45,39]
d: [2,91,1,2,1,2,21,90,71]
f: 122,113,95,91,88,85,33,28,29,30,31,32
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t23_gate_1
% SZS output start Proof for gate_1__t23_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (~v1_xboole_0(k17_gate_1(X0,X1,X2,X3)) <=> (v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [input t23_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (~v1_xboole_0(k17_gate_1(X0,X1,X2,X3)) <=> (v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k17_gate_1(X0,X1,X2,X3) = k1_xboole_0) & ~(k17_gate_1(X0,X1,X2,X3) != k1_gate_1(k1_xboole_0) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [input d17_gate_1]
14. ~! [X0,X1,X2,X3] : (~v1_xboole_0(k17_gate_1(X0,X1,X2,X3)) <=> (v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [flattening 2]
15. ! [X0,X1,X2,X3] : ((~(v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k17_gate_1(X0,X1,X2,X3) = k1_xboole_0) & ~(k17_gate_1(X0,X1,X2,X3) != k1_gate_1(k1_xboole_0) & v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [flattening 13]
16. ? [X0,X1,X2,X3] : (~v1_xboole_0(k17_gate_1(X0,X1,X2,X3)) <~> (v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2,X3] : ((k17_gate_1(X0,X1,X2,X3) = k1_xboole_0 | (v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & (k17_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 15]
20. ? [X0,X1,X2,X3] : (((~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) | v1_xboole_0(k17_gate_1(X0,X1,X2,X3))) & ((v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) | ~v1_xboole_0(k17_gate_1(X0,X1,X2,X3)))) [nnf transformation 16]
21. ? [X0,X1,X2,X3] : ((~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) | v1_xboole_0(k17_gate_1(X0,X1,X2,X3))) & ((v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) | ~v1_xboole_0(k17_gate_1(X0,X1,X2,X3)))) [flattening 20]
22. ? [X0,X1,X2,X3] : ((~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) | v1_xboole_0(k17_gate_1(X0,X1,X2,X3))) & ((v1_xboole_0(X3) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) | ~v1_xboole_0(k17_gate_1(X0,X1,X2,X3)))) => ((~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3))) & ((v1_xboole_0(sK3) & v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0)) | ~v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3)))) [choice axiom]
23. (~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3))) & ((v1_xboole_0(sK3) & v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0)) | ~v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3))) [skolemisation 21,22]
28. v1_xboole_0(sK0) | ~v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. v1_xboole_0(sK1) | ~v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. v1_xboole_0(sK2) | ~v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. v1_xboole_0(sK3) | ~v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
35. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
36. k17_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X3) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:16,wCS:6400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
37. k17_gate_1(X0,X1,X2,X3) = k1_xboole_0 | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. k17_gate_1(X0,X1,X2,X3) = k1_xboole_0 | v1_xboole_0(X1) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. k17_gate_1(X0,X1,X2,X3) = k1_xboole_0 | v1_xboole_0(X2) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. k17_gate_1(X0,X1,X2,X3) = k1_xboole_0 | v1_xboole_0(X3) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. 1 <=> v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3)) [avatar definition]
45. ~v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3)) <- (~1) [avatar component clause 44] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
46. v1_xboole_0(k17_gate_1(sK0,sK1,sK2,sK3)) <- (1) [avatar component clause 44] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
48. 2 <=> v1_xboole_0(sK0) [avatar definition]
49. v1_xboole_0(sK0) <- (2) [avatar component clause 48] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
52. 3 <=> v1_xboole_0(sK1) [avatar definition]
53. v1_xboole_0(sK1) <- (3) [avatar component clause 52] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. 4 <=> v1_xboole_0(sK2) [avatar definition]
57. v1_xboole_0(sK2) <- (4) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 5 <=> v1_xboole_0(sK3) [avatar definition]
61. v1_xboole_0(sK3) <- (5) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. ~v1_xboole_0(sK3) <- (~5) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. 1 | ~2 | ~3 | ~4 | ~5 [avatar split clause 32,60,56,52,48,44]
64. ~1 | 5 [avatar split clause 31,60,44]
65. ~1 | 4 [avatar split clause 30,56,44]
66. ~1 | 3 [avatar split clause 29,52,44]
67. ~1 | 2 [avatar split clause 28,48,44]
69. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
71. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 69] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
72. 6 [avatar split clause 33,69]
84. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK0) <- (~1) [superposition 45,37] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
85. v1_xboole_0(sK0) <- (~1, 6) [subsumption resolution 84,71] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. 2 | 1 | ~6 [avatar split clause 85,69,44,48]
87. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK1) <- (~1) [superposition 45,38] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
88. v1_xboole_0(sK1) <- (~1, 6) [subsumption resolution 87,71] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. 3 | 1 | ~6 [avatar split clause 88,69,44,52]
90. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK2) <- (~1) [superposition 45,39] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. v1_xboole_0(sK2) <- (~1, 6) [subsumption resolution 90,71] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. 4 | 1 | ~6 [avatar split clause 91,69,44,56]
93. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK3) <- (~1) [superposition 45,40] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
94. v1_xboole_0(sK3) <- (~1, 6) [subsumption resolution 93,71] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. $false <- (~1, ~5, 6) [subsumption resolution 94,62] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. 1 | 5 | ~6 [avatar contradiction clause 95]
109. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1) [superposition 46,36] {a:1,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, 5) [subsumption resolution 109,61] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
111. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, 4, 5) [subsumption resolution 110,57] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
112. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK0) <- (1, 3, 4, 5) [subsumption resolution 111,53] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (1, 2, 3, 4, 5) [subsumption resolution 112,49] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
115. 9 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
117. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (9) [avatar component clause 115] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. 9 | ~1 | ~2 | ~3 | ~4 | ~5 [avatar split clause 113,60,56,52,48,44,115]
120. ~v1_xboole_0(k1_xboole_0) <- (9) [resolution 117,35] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
122. $false <- (6, 9) [subsumption resolution 120,71] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
123. ~6 | ~9 [avatar contradiction clause 122]
124. $false [avatar sat refutation 63,64,65,66,67,72,86,89,92,96,118,123] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:70,thDist:-70}
% SZS output end Proof for gate_1__t23_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 56
% Active clauses: 18
% Passive clauses: 18
% Final active clauses: 16
% Main loop iterations started: 18
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 8
% Backward superposition: 9
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 13
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 3
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

