i: [1,110,0,2,1,1,0,0] t19_diraf
a: [3,174,0,2,1,110]
s: 174
i: [1,109,0,2,1,1,0,0] t19_diraf
a: [3,179,0,2,1,109]
s: 179
i: [1,108,0,2,1,1,0,0] t19_diraf
a: [3,184,0,2,1,108]
s: 184
i: [1,112,0,4,1,1,0,0] t19_diraf
a: [3,164,0,4,1,112]
s: 164
i: [1,111,0,4,1,1,0,0] t19_diraf
a: [3,169,0,4,1,111]
s: 169
i: [1,113,0,12,2,1,0,0] t19_diraf
a: [3,155,0,6,1,113]
s: 155
i: [1,114,0,2,1,0,0,255] spc1_boole
a: [3,189,0,2,1,114]
s: 189
i: [1,115,0,2,1,0,0,255] fc1_xboole_0
a: [3,194,0,2,1,115]
s: 194
i: [1,148,0,2,1,0,0,2] rc2_xboole_0
a: [3,199,0,2,1,148]
s: 199
i: [1,149,0,2,1,0,0,1] existence_l1_analoaf
a: [3,204,0,2,1,149]
s: 204
i: [1,150,0,2,1,0,0,2] existence_l1_struct_0
a: [3,209,0,2,1,150]
s: 209
i: [1,151,0,2,1,0,0,2] rc1_xboole_0
a: [3,214,0,2,1,151]
s: 214
i: [1,116,0,4,2,0,0,1] dt_l1_analoaf
s: 116
d: [2,216,1,2,1,1,40,116,174]
a: [3,221,1,2,1,216]
s: 221
i: [1,146,0,4,1,0,0,1] existence_m1_subset_1
s: 146
i: [1,131,0,5,2,0,0,255] t6_boole
s: 131
d: [2,217,1,2,1,1,40,116,204]
a: [3,226,1,2,1,217]
s: 226
i: [1,145,0,5,2,0,0,255] t7_boole
s: 145
i: [1,117,0,6,3,0,0,255] cc7_struct_0
s: 117
i: [1,118,0,6,3,0,0,255] cc5_struct_0
s: 118
i: [1,120,0,6,3,0,0,1] cc17_struct_0
s: 120
d: [2,230,1,4,2,1,40,117,209]
a: [3,236,1,2,1,230]
s: 236
i: [1,121,0,6,3,0,0,1] cc2_struct_0
s: 121
d: [2,232,2,4,2,1,40,117,226]
a: [3,245,2,2,1,232]
s: 245
i: [1,122,0,6,3,0,0,1] cc1_struct_0
s: 122
i: [1,125,0,6,3,0,0,255] cc4_struct_0
s: 125
i: [1,126,0,6,3,0,0,255] cc6_struct_0
s: 126
i: [1,142,0,6,2,0,0,255] t1_subset
s: 142
i: [1,147,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 147
i: [1,123,0,7,3,0,0,255] cc8_struct_0
s: 123
i: [1,127,0,7,3,0,0,255] cc11_struct_0
s: 127
i: [1,128,0,7,3,0,0,255] cc11_struct_0
s: 128
i: [1,129,0,7,3,0,0,255] cc9_struct_0
s: 129
i: [1,132,0,7,3,0,0,255] fc9_struct_0
s: 132
i: [1,136,0,7,3,0,0,1] fc2_struct_0
s: 136
i: [1,138,0,7,3,0,0,255] fc6_struct_0
s: 138
i: [1,139,0,7,3,0,0,255] fc8_struct_0
s: 139
i: [1,140,0,7,3,0,0,1] fc1_struct_0
s: 140
i: [1,141,0,7,3,0,0,255] fc7_struct_0
s: 141
i: [1,144,0,7,3,0,0,255] t8_boole
s: 144
i: [1,143,0,8,3,0,0,255] t2_subset
s: 143
d: [2,261,1,7,2,1,40,143,164]
a: [3,271,1,3,1,261]
s: 271
r: 271
i: [1,130,0,9,4,0,0,255] cc10_struct_0
s: 130
a: [3,270,1,3,1,261]
s: 270
i: [1,137,0,20,6,0,0,1] t1_diraf
s: 137
i: [1,134,0,32,8,0,0,1] d4_diraf
s: 134
d: [2,278,2,4,2,1,40,271,136]
d: [2,279,2,2,1,2,21,278,221]
a: [3,282,2,2,1,279]
s: 282
i: [1,135,0,32,8,0,0,1] d4_diraf
s: 135
a: [3,267,1,4,1,261]
s: 267
i: [1,133,0,38,9,0,0,1] d4_diraf
s: 133
d: [2,262,1,7,2,1,40,143,169]
a: [3,276,1,4,1,262]
s: 276
d: [2,296,2,4,1,1,40,267,147]
a: [3,302,2,4,1,296]
s: 302
d: [2,263,1,6,2,0,40,143,146]
s: 263
d: [2,304,2,4,1,1,40,276,147]
a: [3,310,2,4,1,304]
s: 310
d: [2,228,1,7,3,0,44,131,131]
s: 228
d: [2,251,1,9,4,0,44,123,131]
s: 251
d: [2,312,2,6,2,0,40,263,147]
s: 312
d: [2,253,1,9,4,0,44,129,131]
s: 253
d: [2,292,1,40,12,0,40,134,137]
d: [2,293,1,22,7,0,18,292]
s: 293
d: [2,294,1,40,12,0,40,135,137]
d: [2,295,1,22,7,0,18,294]
s: 295
d: [2,329,3,0,0,7,52,184,179,282,174,164,169,155,295]
e: 329
d: [2,330,2,16,6,1,40,295,155]
d: [2,334,2,12,5,2,21,330,169]
d: [2,335,2,8,4,3,21,334,164]
d: [2,336,2,6,3,4,21,335,174]
d: [2,337,2,4,2,5,21,336,282]
d: [2,338,2,2,1,6,21,337,179]
d: [2,339,2,0,0,7,21,338,184]
e: 339
r: 155
a: [3,159,0,6,1,113]
s: 159
d: [2,341,3,0,0,7,52,184,179,282,174,169,164,159,293]
e: 341
d: [2,342,2,16,6,1,40,159,293]
d: [2,344,2,12,5,2,21,342,164]
d: [2,345,2,8,4,3,21,344,169]
d: [2,346,2,6,3,4,21,345,174]
d: [2,347,2,4,2,5,21,346,282]
d: [2,348,2,2,1,6,21,347,179]
d: [2,349,2,0,0,7,21,348,184]
e: 349
r: 270
r: 282
r: 267
r: 276
a: [3,283,2,2,1,279]
s: 283
d: [2,356,3,0,0,3,52,221,184,283,120]
e: 356
d: [2,357,3,4,2,1,40,283,120]
d: [2,360,3,2,1,2,21,357,184]
d: [2,361,3,0,0,3,21,360,221]
e: 361
d: [2,353,2,4,2,5,21,346,179]
d: [2,354,2,2,1,6,21,353,184]
f: 361,354,339,216,108,109,110,111,112,113
% Refutation found. Thanks to Tanya!
% SZS status Theorem for diraf__t19_diraf
% SZS output start Proof for diraf__t19_diraf
1. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X1,X2) & r2_diraf(X0,X1,X2,X2,X1))))) [input t19_diraf]
2. ~! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X1,X2) & r2_diraf(X0,X1,X2,X2,X1))))) [negated conjecture 1]
6. ! [X0] : (l1_struct_0(X0) => (~v7_struct_0(X0) => ~v2_struct_0(X0))) [input cc17_struct_0]
15. ! [X0] : ((l1_analoaf(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4)))))))) [input d4_diraf]
17. ! [X0] : (l1_analoaf(X0) => l1_struct_0(X0)) [input dt_l1_analoaf]
34. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => r2_analoaf(X0,X1,X2,X1,X2)))) [input t1_diraf]
44. ? [X0] : (? [X1] : (? [X2] : ((~r2_diraf(X0,X1,X2,X1,X2) | ~r2_diraf(X0,X1,X2,X2,X1)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
45. ? [X0] : (? [X1] : (? [X2] : ((~r2_diraf(X0,X1,X2,X1,X2) | ~r2_diraf(X0,X1,X2,X2,X1)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) [flattening 44]
46. ! [X0] : (l1_struct_0(X0) | ~l1_analoaf(X0)) [ennf transformation 17]
51. ! [X0] : ((~v2_struct_0(X0) | v7_struct_0(X0)) | ~l1_struct_0(X0)) [ennf transformation 6]
52. ! [X0] : (~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0)) [flattening 51]
72. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | v2_struct_0(X0))) [ennf transformation 15]
73. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 72]
76. ! [X0] : (! [X1] : (! [X2] : (r2_analoaf(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0))) [ennf transformation 34]
77. ! [X0] : (! [X1] : (! [X2] : (r2_analoaf(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0)) [flattening 76]
92. ? [X0] : (? [X1] : (? [X2] : ((~r2_diraf(X0,X1,X2,X1,X2) | ~r2_diraf(X0,X1,X2,X2,X1)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : ((~r2_diraf(sK0,X1,X2,X1,X2) | ~r2_diraf(sK0,X1,X2,X2,X1)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
93. ? [X1] : (? [X2] : ((~r2_diraf(sK0,X1,X2,X1,X2) | ~r2_diraf(sK0,X1,X2,X2,X1)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : ((~r2_diraf(sK0,sK1,X2,sK1,X2) | ~r2_diraf(sK0,sK1,X2,X2,sK1)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
94. ? [X2] : ((~r2_diraf(sK0,sK1,X2,sK1,X2) | ~r2_diraf(sK0,sK1,X2,X2,sK1)) & m1_subset_1(X2,u1_struct_0(sK0))) => ((~r2_diraf(sK0,sK1,sK2,sK1,sK2) | ~r2_diraf(sK0,sK1,sK2,sK2,sK1)) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
95. (((~r2_diraf(sK0,sK1,sK2,sK1,sK2) | ~r2_diraf(sK0,sK1,sK2,sK2,sK1)) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0) [skolemisation 45,94,93,92]
96. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r2_diraf(X0,X1,X2,X3,X4) | (~r2_analoaf(X0,X1,X2,X4,X3) & ~r2_analoaf(X0,X1,X2,X3,X4))) & ((r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4)) | ~r2_diraf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [nnf transformation 73]
97. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r2_diraf(X0,X1,X2,X3,X4) | (~r2_analoaf(X0,X1,X2,X4,X3) & ~r2_analoaf(X0,X1,X2,X3,X4))) & (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4) | ~r2_diraf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 96]
108. ~v7_struct_0(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
109. v2_analoaf(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
110. l1_analoaf(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
111. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 95] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
112. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 95] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
113. ~r2_diraf(sK0,sK1,sK2,sK1,sK2) | ~r2_diraf(sK0,sK1,sK2,sK2,sK1) [cnf transformation 95] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
116. ~l1_analoaf(X0) | l1_struct_0(X0) [cnf transformation 46] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. ~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 52] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
134. ~r2_analoaf(X0,X1,X2,X3,X4) | r2_diraf(X0,X1,X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 97] {a:0,w:32,wCS:12800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
135. ~r2_analoaf(X0,X1,X2,X4,X3) | r2_diraf(X0,X1,X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 97] {a:0,w:32,wCS:12800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
137. r2_analoaf(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 77] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
153. 1 <=> r2_diraf(sK0,sK1,sK2,sK2,sK1) [avatar definition]
155. ~r2_diraf(sK0,sK1,sK2,sK2,sK1) <- (~1) [avatar component clause 153] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
157. 2 <=> r2_diraf(sK0,sK1,sK2,sK1,sK2) [avatar definition]
159. ~r2_diraf(sK0,sK1,sK2,sK1,sK2) <- (~2) [avatar component clause 157] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
160. ~1 | ~2 [avatar split clause 113,157,153]
162. 3 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
164. m1_subset_1(sK2,u1_struct_0(sK0)) <- (3) [avatar component clause 162] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
165. 3 [avatar split clause 112,162]
167. 4 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
169. m1_subset_1(sK1,u1_struct_0(sK0)) <- (4) [avatar component clause 167] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
170. 4 [avatar split clause 111,167]
172. 5 <=> l1_analoaf(sK0) [avatar definition]
174. l1_analoaf(sK0) <- (5) [avatar component clause 172] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
175. 5 [avatar split clause 110,172]
177. 6 <=> v2_analoaf(sK0) [avatar definition]
179. v2_analoaf(sK0) <- (6) [avatar component clause 177] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
180. 6 [avatar split clause 109,177]
182. 7 <=> v7_struct_0(sK0) [avatar definition]
184. ~v7_struct_0(sK0) <- (~7) [avatar component clause 182] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
185. ~7 [avatar split clause 108,182]
216. l1_struct_0(sK0) <- (5) [resolution 116,174] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
219. 14 <=> l1_struct_0(sK0) [avatar definition]
221. l1_struct_0(sK0) <- (14) [avatar component clause 219] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
222. 14 | ~5 [avatar split clause 216,172,219]
281. 23 <=> v2_struct_0(sK0) [avatar definition]
282. ~v2_struct_0(sK0) <- (~23) [avatar component clause 281] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
283. v2_struct_0(sK0) <- (23) [avatar component clause 281] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
292. r2_diraf(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [resolution 134,137] {a:1,w:40,wCS:16000,sine:1,thAx:0,allAx:2,thDist:-2}
293. r2_diraf(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [duplicate literal removal 292] {a:1,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
294. r2_diraf(X0,X1,X2,X2,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [resolution 135,137] {a:1,w:40,wCS:16000,sine:1,thAx:0,allAx:2,thDist:-2}
295. r2_diraf(X0,X1,X2,X2,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [duplicate literal removal 294] {a:1,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
330. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1) [resolution 295,155] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
334. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 4) [subsumption resolution 330,169] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
335. ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 3, 4) [subsumption resolution 334,164] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
336. v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 3, 4, 5) [subsumption resolution 335,174] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
337. ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 3, 4, 5, ~23) [subsumption resolution 336,282] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
338. v7_struct_0(sK0) <- (~1, 3, 4, 5, 6, ~23) [subsumption resolution 337,179] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
339. $false <- (~1, 3, 4, 5, 6, ~7, ~23) [subsumption resolution 338,184] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
340. 1 | ~3 | ~4 | ~5 | ~6 | 7 | 23 [avatar contradiction clause 339]
342. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2) [resolution 159,293] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
344. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 3) [subsumption resolution 342,164] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
345. ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 3, 4) [subsumption resolution 344,169] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
346. v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 3, 4, 5) [subsumption resolution 345,174] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
353. v2_struct_0(sK0) | v7_struct_0(sK0) <- (~2, 3, 4, 5, 6) [subsumption resolution 346,179] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
354. v2_struct_0(sK0) <- (~2, 3, 4, 5, 6, ~7) [subsumption resolution 353,184] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
355. 23 | 2 | ~3 | ~4 | ~5 | ~6 | 7 [avatar split clause 354,182,177,172,167,162,157,281]
357. v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (23) [resolution 283,120] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
360. ~l1_struct_0(sK0) <- (~7, 23) [subsumption resolution 357,184] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
361. $false <- (~7, 14, 23) [subsumption resolution 360,221] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
362. 7 | ~14 | ~23 [avatar contradiction clause 361]
364. $false [avatar sat refutation 160,165,170,175,180,185,222,340,355,362] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:250,thDist:-250}
% SZS output end Proof for diraf__t19_diraf
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 44
% Generated clauses: 156
% Active clauses: 57
% Passive clauses: 65
% Final active clauses: 51
% Final passive clauses: 3
% Main loop iterations started: 62
% 
% >>> Simplifying Inferences
% Duplicate literals: 37
% Fw subsumption resolutions: 19
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 16
% 
% >>> Generating Inferences
% Binary resolution: 39
% Unit resulting resolution: 3
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 33
% SAT solver unit clauses: 11
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% non unit clause index maintenance: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

