i: [1,27,0,3,1,1,0,0] t18_arytm_1
a: [3,45,0,3,1,27]
s: 45
i: [1,28,0,5,1,1,0,0] t18_arytm_1
a: [3,40,0,5,1,28]
s: 40
i: [1,29,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,50,0,3,1,29]
s: 50
i: [1,30,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,55,0,3,1,30]
s: 55
i: [1,34,0,4,1,0,0,1] existence_m1_subset_1
s: 34
i: [1,31,0,8,2,0,0,255] l13_arytm_1
s: 31
i: [1,35,0,11,3,0,0,255] dt_k1_arytm_1
s: 35
i: [1,36,0,12,4,0,0,255] connectedness_r1_arytm_2
s: 36
d: [2,68,1,9,3,1,40,36,45]
s: 68
i: [1,32,0,16,4,0,0,255] d2_arytm_1
s: 32
i: [1,33,0,18,4,0,0,255] d2_arytm_1
s: 33
d: [2,71,2,6,2,1,40,68,45]
d: [2,73,2,3,1,1,18,71]
a: [3,77,2,3,1,73]
s: 77
d: [2,57,1,12,4,0,44,35,31]
d: [2,58,1,6,2,0,18,57]
a: [3,65,1,3,1,58]
s: 65
r: 65
d: [2,72,2,8,2,1,40,68,34]
a: [3,82,2,4,1,72]
s: 82
a: [3,61,1,3,1,58]
s: 61
r: 68
r: 36
r: 35
r: 33
r: 32
r: 31
d: [2,114,2,0,0,1,40,61,34]
e: 114
d: [2,113,2,0,0,2,40,61,45]
e: 113
d: [2,112,2,0,0,2,52,45,61]
e: 112
r: 61
s: 31
d: [2,101,2,6,2,2,40,65,68]
a: [3,106,2,3,1,101]
s: 106
s: 35
s: 68
s: 36
s: 32
s: 65
s: 33
d: [2,70,2,16,4,1,40,68,35]
s: 70
d: [2,88,1,17,5,0,44,32,31]
d: [2,93,1,11,3,0,18,88]
d: [2,94,1,8,2,0,21,93,36]
s: 94
d: [2,147,2,0,0,2,52,45,40,94]
e: 147
d: [2,148,2,6,2,1,44,40,94]
d: [2,149,2,3,1,1,19,148]
d: [2,151,2,0,0,2,21,149,45]
e: 151
f: 151,27,28
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t18_arytm_1
% SZS output start Proof for arytm_1__t18_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => k2_arytm_1(X0,X0) = k11_arytm_3) [input t18_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => k2_arytm_1(X0,X0) = k11_arytm_3) [negated conjecture 1]
3. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => (r1_arytm_2(X1,X0) | r1_arytm_2(X0,X1))) [input connectedness_r1_arytm_2]
4. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((~r1_arytm_2(X1,X0) => k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,k1_arytm_1(X1,X0))) & (r1_arytm_2(X1,X0) => k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1))))) [input d2_arytm_1]
14. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => k11_arytm_3 = k1_arytm_1(X0,X0)) [input l13_arytm_1]
16. ? [X0] : (k2_arytm_1(X0,X0) != k11_arytm_3 & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
17. ! [X0] : (k11_arytm_3 = k1_arytm_1(X0,X0) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 14]
18. ! [X0] : (! [X1] : (((k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,k1_arytm_1(X1,X0)) | r1_arytm_2(X1,X0)) & (k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 4]
21. ! [X0,X1] : ((r1_arytm_2(X1,X0) | r1_arytm_2(X0,X1)) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 3]
22. ! [X0,X1] : (r1_arytm_2(X1,X0) | r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 21]
23. ? [X0] : (k2_arytm_1(X0,X0) != k11_arytm_3 & m1_subset_1(X0,k2_arytm_2)) => (k11_arytm_3 != k2_arytm_1(sK0,sK0) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
24. k11_arytm_3 != k2_arytm_1(sK0,sK0) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 16,23]
27. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 24] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. k11_arytm_3 != k2_arytm_1(sK0,sK0) [cnf transformation 24] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. k11_arytm_3 = k1_arytm_1(X0,X0) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 17] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
32. k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1) | ~r1_arytm_2(X1,X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 18] {a:0,w:16,wCS:6400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
36. ~m1_subset_1(X1,k2_arytm_2) | r1_arytm_2(X0,X1) | r1_arytm_2(X1,X0) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 22] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. 1 <=> k11_arytm_3 = k2_arytm_1(sK0,sK0) [avatar definition]
40. k11_arytm_3 != k2_arytm_1(sK0,sK0) <- (~1) [avatar component clause 38] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
41. ~1 [avatar split clause 28,38]
43. 2 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
45. m1_subset_1(sK0,k2_arytm_2) <- (2) [avatar component clause 43] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
46. 2 [avatar split clause 27,43]
88. k2_arytm_1(X0,X0) = k11_arytm_3 | ~r1_arytm_2(X0,X0) | ~m1_subset_1(X0,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [superposition 32,31] {a:1,w:17,wCS:6800,sine:255,thAx:0,allAx:2,thDist:-2}
93. k2_arytm_1(X0,X0) = k11_arytm_3 | ~r1_arytm_2(X0,X0) | ~m1_subset_1(X0,k2_arytm_2) [duplicate literal removal 88] {a:1,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
94. k2_arytm_1(X0,X0) = k11_arytm_3 | ~m1_subset_1(X0,k2_arytm_2) [subsumption resolution 93,36] {a:1,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
148. k11_arytm_3 != k11_arytm_3 | ~m1_subset_1(sK0,k2_arytm_2) <- (~1) [superposition 40,94] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
149. ~m1_subset_1(sK0,k2_arytm_2) <- (~1) [trivial inequality removal 148] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
151. $false <- (~1, 2) [subsumption resolution 149,45] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
152. 1 | ~2 [avatar contradiction clause 151]
153. $false [avatar sat refutation 41,46,152] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:64,thDist:-64}
% SZS output end Proof for arytm_1__t18_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 10
% Generated clauses: 86
% Active clauses: 25
% Passive clauses: 43
% Final active clauses: 17
% Final passive clauses: 16
% Main loop iterations started: 25
% 
% >>> Simplifying Inferences
% Duplicate literals: 30
% Trivial inequalities: 1
% Fw subsumption resolutions: 7
% 
% >>> Deletion Inferences
% Forward subsumptions: 5
% Backward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 18
% Unit resulting resolution: 2
% Forward superposition: 10
% Backward superposition: 5
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 5
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

