i: [1,12,0,8,1,1,0,0] t8_borsuk_5
a: [3,59,0,8,1,12]
s: 59
i: [1,11,0,8,1,1,0,0] t8_borsuk_5
a: [3,64,0,8,1,11]
s: 64
i: [1,35,0,8,1,0,0,255] d9_zfmisc_1
s: 35
i: [1,36,0,8,1,0,0,255] d9_zfmisc_1
s: 36
i: [1,37,0,8,1,0,0,255] d9_zfmisc_1
s: 37
i: [1,38,0,8,1,0,0,255] d9_zfmisc_1
s: 38
i: [1,39,0,8,1,0,0,255] d9_zfmisc_1
s: 39
i: [1,40,0,8,1,0,0,255] d9_zfmisc_1
s: 40
i: [1,41,0,8,1,0,0,255] d9_zfmisc_1
s: 41
i: [1,42,0,8,1,0,0,255] d9_zfmisc_1
s: 42
i: [1,43,0,8,1,0,0,255] d9_zfmisc_1
s: 43
i: [1,44,0,8,1,0,0,255] d9_zfmisc_1
s: 44
i: [1,45,0,8,1,0,0,255] d9_zfmisc_1
s: 45
i: [1,46,0,8,1,0,0,255] d9_zfmisc_1
s: 46
i: [1,47,0,8,1,0,0,255] d9_zfmisc_1
s: 47
i: [1,48,0,8,1,0,0,255] d9_zfmisc_1
s: 48
i: [1,49,0,8,1,0,0,255] d9_zfmisc_1
s: 49
i: [1,50,0,8,1,0,0,255] d9_zfmisc_1
s: 50
i: [1,51,0,8,1,0,0,255] d9_zfmisc_1
s: 51
i: [1,52,0,8,1,0,0,255] d9_zfmisc_1
s: 52
i: [1,53,0,8,1,0,0,255] d9_zfmisc_1
s: 53
i: [1,54,0,8,1,0,0,255] d9_zfmisc_1
s: 54
i: [1,55,0,8,1,0,0,255] d9_zfmisc_1
s: 55
i: [1,34,0,71,22,0,0,255] d9_zfmisc_1
s: 34
d: [2,66,1,63,21,1,40,34,59]
a: [3,116,1,3,1,66]
s: 116
r: 64
r: 59
d: [2,194,0,8,1,2,23,64,116]
d: [2,196,0,0,0,2,21,194,45]
e: 196
r: 116
s: 59
a: [3,115,1,3,1,66]
s: 115
s: 64
a: [3,120,1,3,1,66]
s: 120
r: 64
r: 59
d: [2,200,0,8,1,2,23,64,120]
d: [2,203,0,0,0,2,21,200,41]
e: 203
r: 120
s: 59
s: 64
a: [3,119,1,3,1,66]
s: 119
a: [3,124,1,3,1,66]
s: 124
r: 64
r: 59
d: [2,208,0,8,1,2,23,64,124]
d: [2,212,0,0,0,2,21,208,38]
e: 212
r: 124
a: [3,123,1,3,1,66]
s: 123
s: 59
s: 64
a: [3,132,1,3,1,66]
s: 132
r: 64
r: 59
d: [2,219,0,8,1,2,23,64,132]
d: [2,224,0,0,0,2,21,219,35]
e: 224
r: 132
s: 59
a: [3,131,1,3,1,66]
s: 131
s: 64
a: [3,136,1,3,1,66]
s: 136
r: 64
r: 59
r: 115
d: [2,231,0,8,1,2,23,64,136]
d: [2,234,0,0,0,2,21,231,55]
e: 234
r: 136
s: 115
s: 59
a: [3,135,1,3,1,66]
s: 135
s: 64
a: [3,144,1,3,1,66]
s: 144
r: 64
r: 59
r: 123
d: [2,242,0,8,1,2,23,64,144]
d: [2,245,0,0,0,2,21,242,53]
e: 245
r: 144
s: 123
s: 59
a: [3,128,1,3,1,66]
s: 128
r: 59
r: 131
s: 64
d: [2,263,0,8,1,2,22,64,128]
d: [2,264,0,0,0,2,21,263,36]
e: 264
r: 128
s: 59
a: [3,127,1,3,1,66]
s: 127
s: 64
s: 131
a: [3,143,1,3,1,66]
s: 143
a: [3,152,1,3,1,66]
s: 152
r: 131
r: 64
r: 59
d: [2,273,0,8,1,2,23,64,152]
d: [2,276,0,0,0,2,21,273,51]
e: 276
r: 152
s: 131
s: 59
a: [3,151,1,3,1,66]
s: 151
s: 64
a: [3,160,1,3,1,66]
s: 160
r: 64
r: 59
r: 123
d: [2,287,0,8,1,2,23,64,160]
d: [2,291,0,0,0,2,21,287,48]
e: 291
r: 160
s: 123
s: 59
a: [3,159,1,3,1,66]
s: 159
s: 64
a: [3,163,1,3,1,66]
s: 163
a: [3,172,1,3,1,66]
s: 172
r: 64
r: 59
r: 123
d: [2,303,0,8,1,2,23,64,172]
d: [2,308,0,0,0,2,21,303,44]
e: 308
r: 172
s: 59
s: 123
s: 64
a: [3,148,1,3,1,66]
s: 148
r: 64
r: 59
r: 163
r: 127
d: [2,321,0,8,1,2,23,64,148]
d: [2,325,0,0,0,2,21,321,52]
e: 325
r: 148
s: 127
s: 59
a: [3,147,1,3,1,66]
s: 147
s: 64
a: [3,164,1,3,1,66]
s: 164
r: 64
r: 59
r: 127
d: [2,347,0,8,1,2,23,64,164]
d: [2,351,0,0,0,2,21,347,47]
e: 351
r: 164
s: 127
s: 59
s: 163
s: 64
a: [3,171,1,3,1,66]
s: 171
a: [3,175,1,3,1,66]
s: 175
a: [3,180,1,3,1,66]
s: 180
r: 64
r: 59
r: 131
d: [2,365,0,8,1,2,23,64,180]
d: [2,369,0,0,0,2,21,365,42]
e: 369
r: 180
s: 131
s: 59
s: 64
a: [3,179,1,3,1,66]
s: 179
a: [3,183,1,3,1,66]
s: 183
a: [3,188,1,3,1,66]
s: 188
r: 64
r: 59
r: 131
d: [2,386,0,8,1,2,23,64,188]
d: [2,391,0,0,0,2,21,386,39]
e: 391
r: 188
s: 59
s: 64
a: [3,112,1,3,1,66]
s: 112
r: 59
r: 64
r: 115
r: 119
r: 123
r: 127
d: [2,410,0,8,1,2,23,64,112]
d: [2,415,0,0,0,2,21,410,50]
e: 415
r: 112
s: 59
a: [3,111,1,3,1,66]
s: 111
s: 64
s: 115
s: 119
s: 123
s: 127
s: 131
a: [3,156,1,3,1,66]
s: 156
r: 119
r: 64
r: 59
r: 179
r: 175
r: 171
d: [2,430,0,8,1,2,23,64,156]
d: [2,435,0,0,0,2,21,430,49]
e: 435
r: 156
s: 119
s: 59
a: [3,155,1,3,1,66]
s: 155
s: 64
s: 171
a: [3,176,1,3,1,66]
s: 176
r: 59
r: 64
r: 127
r: 183
d: [2,490,0,8,1,2,23,64,176]
d: [2,495,0,0,0,2,21,490,43]
e: 495
r: 176
s: 59
s: 127
s: 64
a: [3,140,1,3,1,66]
s: 140
r: 64
r: 59
r: 171
r: 155
r: 119
d: [2,512,0,8,1,2,23,64,140]
d: [2,517,0,0,0,2,21,512,54]
e: 517
r: 140
s: 119
s: 59
a: [3,139,1,3,1,66]
s: 139
s: 64
s: 155
s: 171
s: 175
s: 179
a: [3,184,1,3,1,66]
s: 184
r: 59
r: 64
r: 127
d: [2,582,0,8,1,2,23,64,184]
d: [2,587,0,0,0,2,21,582,40]
e: 587
r: 184
s: 59
s: 127
s: 64
s: 183
a: [3,187,1,3,1,66]
s: 187
a: [3,192,1,3,1,66]
s: 192
r: 64
r: 59
r: 131
d: [2,607,0,8,1,2,23,64,192]
d: [2,613,0,0,0,2,21,607,37]
e: 613
r: 192
s: 131
s: 59
a: [3,168,1,3,1,66]
s: 168
r: 59
r: 131
r: 187
r: 179
s: 64
d: [2,641,0,8,1,2,22,64,168]
d: [2,642,0,0,0,2,21,641,46]
e: 642
d: [2,615,1,63,21,1,40,59,34]
d: [2,616,1,60,20,2,21,615,187]
d: [2,617,1,57,19,3,21,616,183]
d: [2,618,1,54,18,4,21,617,179]
d: [2,619,1,51,17,5,21,618,175]
d: [2,620,1,48,16,6,21,619,171]
d: [2,621,1,45,15,7,21,620,163]
d: [2,622,1,42,14,8,21,621,159]
d: [2,623,1,39,13,9,21,622,155]
d: [2,624,1,36,12,10,21,623,151]
d: [2,625,1,33,11,11,21,624,147]
d: [2,626,1,30,10,12,21,625,143]
d: [2,627,1,27,9,13,21,626,139]
d: [2,628,1,24,8,14,21,627,135]
d: [2,629,1,21,7,15,21,628,131]
d: [2,630,1,18,6,16,21,629,127]
d: [2,631,1,15,5,17,21,630,123]
d: [2,632,1,12,4,18,21,631,119]
d: [2,633,1,9,3,19,21,632,115]
d: [2,634,1,6,2,20,21,633,111]
f: 642,634,613,587,517,495,435,415,391,369,351,325,308,291,276,264,245,234,224,212,203,196,11,12
% Refutation found. Thanks to Tanya!
% SZS status Theorem for borsuk_5__t8_borsuk_5
% SZS output start Proof for borsuk_5__t8_borsuk_5
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) => r5_zfmisc_1(X6,X0,X1,X2,X3,X4,X5)) [input t8_borsuk_5]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) => r5_zfmisc_1(X6,X0,X1,X2,X3,X4,X5)) [negated conjecture 1]
3. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) <=> (X5 != X6 & X4 != X6 & X4 != X5 & X3 != X6 & X3 != X5 & X3 != X4 & X2 != X6 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X6 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X6 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1)) [input d9_zfmisc_1]
4. ~! [X0,X1,X2,X3,X4,X5,X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) => r5_zfmisc_1(X6,X0,X1,X2,X3,X4,X5)) [flattening 2]
5. ! [X0,X1,X2,X3,X4,X5,X6] : (r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) <=> (X5 != X6 & X4 != X6 & X4 != X5 & X3 != X6 & X3 != X5 & X3 != X4 & X2 != X6 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X6 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X6 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1)) [flattening 3]
6. ? [X0,X1,X2,X3,X4,X5,X6] : (~r5_zfmisc_1(X6,X0,X1,X2,X3,X4,X5) & r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6)) [ennf transformation 4]
7. ? [X0,X1,X2,X3,X4,X5,X6] : (~r5_zfmisc_1(X6,X0,X1,X2,X3,X4,X5) & r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6)) => (~r5_zfmisc_1(sK6,sK0,sK1,sK2,sK3,sK4,sK5) & r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6)) [choice axiom]
8. ~r5_zfmisc_1(sK6,sK0,sK1,sK2,sK3,sK4,sK5) & r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6) [skolemisation 6,7]
9. ! [X0,X1,X2,X3,X4,X5,X6] : ((r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) | (X5 = X6 | X4 = X6 | X4 = X5 | X3 = X6 | X3 = X5 | X3 = X4 | X2 = X6 | X2 = X5 | X2 = X4 | X2 = X3 | X1 = X6 | X1 = X5 | X1 = X4 | X1 = X3 | X1 = X2 | X0 = X6 | X0 = X5 | X0 = X4 | X0 = X3 | X0 = X2 | X0 = X1)) & ((X5 != X6 & X4 != X6 & X4 != X5 & X3 != X6 & X3 != X5 & X3 != X4 & X2 != X6 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X6 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X6 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1) | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6))) [nnf transformation 5]
10. ! [X0,X1,X2,X3,X4,X5,X6] : ((r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) | X5 = X6 | X4 = X6 | X4 = X5 | X3 = X6 | X3 = X5 | X3 = X4 | X2 = X6 | X2 = X5 | X2 = X4 | X2 = X3 | X1 = X6 | X1 = X5 | X1 = X4 | X1 = X3 | X1 = X2 | X0 = X6 | X0 = X5 | X0 = X4 | X0 = X3 | X0 = X2 | X0 = X1) & ((X5 != X6 & X4 != X6 & X4 != X5 & X3 != X6 & X3 != X5 & X3 != X4 & X2 != X6 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X6 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X6 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1) | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6))) [flattening 9]
11. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6) [cnf transformation 8] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
12. ~r5_zfmisc_1(sK6,sK0,sK1,sK2,sK3,sK4,sK5) [cnf transformation 8] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
13. X0 != X1 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
14. X0 != X2 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
15. X0 != X3 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
16. X0 != X4 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
17. X0 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
18. X0 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
19. X1 != X2 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
20. X1 != X3 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
21. X1 != X4 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
22. X1 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
23. X1 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
24. X2 != X3 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
25. X2 != X4 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
26. X2 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
27. X2 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
28. X3 != X4 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
29. X3 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
30. X3 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
31. X4 != X5 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
32. X4 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
33. X5 != X6 | ~r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) [cnf transformation 10] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
34. r5_zfmisc_1(X0,X1,X2,X3,X4,X5,X6) | X5 = X6 | X4 = X6 | X4 = X5 | X3 = X6 | X3 = X5 | X3 = X4 | X2 = X6 | X2 = X5 | X2 = X4 | X2 = X3 | X1 = X6 | X1 = X5 | X1 = X4 | X1 = X3 | X1 = X2 | X0 = X6 | X0 = X5 | X0 = X4 | X0 = X3 | X0 = X2 | X0 = X1 [cnf transformation 10] {a:0,w:71,wCS:28400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
35. ~r5_zfmisc_1(X0,X1,X2,X3,X4,X6,X6) [equality resolution 33] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
36. ~r5_zfmisc_1(X0,X1,X2,X3,X6,X5,X6) [equality resolution 32] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
37. ~r5_zfmisc_1(X0,X1,X2,X3,X5,X5,X6) [equality resolution 31] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. ~r5_zfmisc_1(X0,X1,X2,X6,X4,X5,X6) [equality resolution 30] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. ~r5_zfmisc_1(X0,X1,X2,X5,X4,X5,X6) [equality resolution 29] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. ~r5_zfmisc_1(X0,X1,X2,X4,X4,X5,X6) [equality resolution 28] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
41. ~r5_zfmisc_1(X0,X1,X6,X3,X4,X5,X6) [equality resolution 27] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. ~r5_zfmisc_1(X0,X1,X5,X3,X4,X5,X6) [equality resolution 26] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. ~r5_zfmisc_1(X0,X1,X4,X3,X4,X5,X6) [equality resolution 25] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. ~r5_zfmisc_1(X0,X1,X3,X3,X4,X5,X6) [equality resolution 24] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
45. ~r5_zfmisc_1(X0,X6,X2,X3,X4,X5,X6) [equality resolution 23] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
46. ~r5_zfmisc_1(X0,X5,X2,X3,X4,X5,X6) [equality resolution 22] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
47. ~r5_zfmisc_1(X0,X4,X2,X3,X4,X5,X6) [equality resolution 21] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
48. ~r5_zfmisc_1(X0,X3,X2,X3,X4,X5,X6) [equality resolution 20] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
49. ~r5_zfmisc_1(X0,X2,X2,X3,X4,X5,X6) [equality resolution 19] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
50. ~r5_zfmisc_1(X6,X1,X2,X3,X4,X5,X6) [equality resolution 18] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
51. ~r5_zfmisc_1(X5,X1,X2,X3,X4,X5,X6) [equality resolution 17] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
52. ~r5_zfmisc_1(X4,X1,X2,X3,X4,X5,X6) [equality resolution 16] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
53. ~r5_zfmisc_1(X3,X1,X2,X3,X4,X5,X6) [equality resolution 15] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
54. ~r5_zfmisc_1(X2,X1,X2,X3,X4,X5,X6) [equality resolution 14] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
55. ~r5_zfmisc_1(X1,X1,X2,X3,X4,X5,X6) [equality resolution 13] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
57. 1 <=> r5_zfmisc_1(sK6,sK0,sK1,sK2,sK3,sK4,sK5) [avatar definition]
59. ~r5_zfmisc_1(sK6,sK0,sK1,sK2,sK3,sK4,sK5) <- (~1) [avatar component clause 57] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. ~1 [avatar split clause 12,57]
62. 2 <=> r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6) [avatar definition]
64. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6) <- (2) [avatar component clause 62] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. 2 [avatar split clause 11,62]
110. 3 <=> sK0 = sK6 [avatar definition]
111. sK0 != sK6 <- (~3) [avatar component clause 110] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
112. sK0 = sK6 <- (3) [avatar component clause 110] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 4 <=> sK1 = sK6 [avatar definition]
115. sK1 != sK6 <- (~4) [avatar component clause 114] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. sK1 = sK6 <- (4) [avatar component clause 114] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. 5 <=> sK2 = sK6 [avatar definition]
119. sK2 != sK6 <- (~5) [avatar component clause 118] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
120. sK2 = sK6 <- (5) [avatar component clause 118] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
122. 6 <=> sK3 = sK6 [avatar definition]
123. sK3 != sK6 <- (~6) [avatar component clause 122] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. sK3 = sK6 <- (6) [avatar component clause 122] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. 7 <=> sK4 = sK6 [avatar definition]
127. sK4 != sK6 <- (~7) [avatar component clause 126] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. sK4 = sK6 <- (7) [avatar component clause 126] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
130. 8 <=> sK5 = sK6 [avatar definition]
131. sK5 != sK6 <- (~8) [avatar component clause 130] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
132. sK5 = sK6 <- (8) [avatar component clause 130] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. 9 <=> sK0 = sK1 [avatar definition]
135. sK0 != sK1 <- (~9) [avatar component clause 134] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
136. sK0 = sK1 <- (9) [avatar component clause 134] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
138. 10 <=> sK0 = sK2 [avatar definition]
139. sK0 != sK2 <- (~10) [avatar component clause 138] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
140. sK0 = sK2 <- (10) [avatar component clause 138] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. 11 <=> sK0 = sK3 [avatar definition]
143. sK0 != sK3 <- (~11) [avatar component clause 142] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. sK0 = sK3 <- (11) [avatar component clause 142] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
146. 12 <=> sK0 = sK4 [avatar definition]
147. sK0 != sK4 <- (~12) [avatar component clause 146] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
148. sK0 = sK4 <- (12) [avatar component clause 146] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
150. 13 <=> sK0 = sK5 [avatar definition]
151. sK0 != sK5 <- (~13) [avatar component clause 150] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
152. sK0 = sK5 <- (13) [avatar component clause 150] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
154. 14 <=> sK1 = sK2 [avatar definition]
155. sK1 != sK2 <- (~14) [avatar component clause 154] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
156. sK1 = sK2 <- (14) [avatar component clause 154] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
158. 15 <=> sK1 = sK3 [avatar definition]
159. sK1 != sK3 <- (~15) [avatar component clause 158] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
160. sK1 = sK3 <- (15) [avatar component clause 158] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
162. 16 <=> sK1 = sK4 [avatar definition]
163. sK1 != sK4 <- (~16) [avatar component clause 162] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
164. sK1 = sK4 <- (16) [avatar component clause 162] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
166. 17 <=> sK1 = sK5 [avatar definition]
168. sK1 = sK5 <- (17) [avatar component clause 166] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
170. 18 <=> sK2 = sK3 [avatar definition]
171. sK2 != sK3 <- (~18) [avatar component clause 170] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
172. sK2 = sK3 <- (18) [avatar component clause 170] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
174. 19 <=> sK2 = sK4 [avatar definition]
175. sK2 != sK4 <- (~19) [avatar component clause 174] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
176. sK2 = sK4 <- (19) [avatar component clause 174] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
178. 20 <=> sK2 = sK5 [avatar definition]
179. sK2 != sK5 <- (~20) [avatar component clause 178] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
180. sK2 = sK5 <- (20) [avatar component clause 178] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
182. 21 <=> sK3 = sK4 [avatar definition]
183. sK3 != sK4 <- (~21) [avatar component clause 182] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
184. sK3 = sK4 <- (21) [avatar component clause 182] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
186. 22 <=> sK3 = sK5 [avatar definition]
187. sK3 != sK5 <- (~22) [avatar component clause 186] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
188. sK3 = sK5 <- (22) [avatar component clause 186] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
190. 23 <=> sK4 = sK5 [avatar definition]
192. sK4 = sK5 <- (23) [avatar component clause 190] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
194. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK1) <- (2, 4) [backward demodulation 64,116] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
196. $false <- (2, 4) [subsumption resolution 194,45] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
197. ~2 | ~4 [avatar contradiction clause 196]
200. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK2) <- (2, 5) [backward demodulation 64,120] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
203. $false <- (2, 5) [subsumption resolution 200,41] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
204. ~2 | ~5 [avatar contradiction clause 203]
208. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK3) <- (2, 6) [backward demodulation 64,124] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
212. $false <- (2, 6) [subsumption resolution 208,38] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
213. ~2 | ~6 [avatar contradiction clause 212]
219. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK5) <- (2, 8) [backward demodulation 64,132] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
224. $false <- (2, 8) [subsumption resolution 219,35] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
225. ~2 | ~8 [avatar contradiction clause 224]
231. r5_zfmisc_1(sK0,sK0,sK2,sK3,sK4,sK5,sK6) <- (2, 9) [backward demodulation 64,136] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
234. $false <- (2, 9) [subsumption resolution 231,55] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
235. ~2 | ~9 [avatar contradiction clause 234]
242. r5_zfmisc_1(sK0,sK1,sK2,sK0,sK4,sK5,sK6) <- (2, 11) [backward demodulation 64,144] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
245. $false <- (2, 11) [subsumption resolution 242,53] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
246. ~2 | ~11 [avatar contradiction clause 245]
263. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK4) <- (2, 7) [forward demodulation 64,128] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
264. $false <- (2, 7) [subsumption resolution 263,36] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
265. ~2 | ~7 [avatar contradiction clause 264]
273. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK0,sK6) <- (2, 13) [backward demodulation 64,152] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
276. $false <- (2, 13) [subsumption resolution 273,51] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
277. ~2 | ~13 [avatar contradiction clause 276]
287. r5_zfmisc_1(sK0,sK1,sK2,sK1,sK4,sK5,sK6) <- (2, 15) [backward demodulation 64,160] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
291. $false <- (2, 15) [subsumption resolution 287,48] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
292. ~2 | ~15 [avatar contradiction clause 291]
303. r5_zfmisc_1(sK0,sK1,sK2,sK2,sK4,sK5,sK6) <- (2, 18) [backward demodulation 64,172] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
308. $false <- (2, 18) [subsumption resolution 303,44] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
309. ~2 | ~18 [avatar contradiction clause 308]
321. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK0,sK5,sK6) <- (2, 12) [backward demodulation 64,148] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
325. $false <- (2, 12) [subsumption resolution 321,52] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
326. ~2 | ~12 [avatar contradiction clause 325]
347. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK1,sK5,sK6) <- (2, 16) [backward demodulation 64,164] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
351. $false <- (2, 16) [subsumption resolution 347,47] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
352. ~2 | ~16 [avatar contradiction clause 351]
365. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK2,sK6) <- (2, 20) [backward demodulation 64,180] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
369. $false <- (2, 20) [subsumption resolution 365,42] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
370. ~2 | ~20 [avatar contradiction clause 369]
386. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK3,sK6) <- (2, 22) [backward demodulation 64,188] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
391. $false <- (2, 22) [subsumption resolution 386,39] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
392. ~2 | ~22 [avatar contradiction clause 391]
410. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5,sK0) <- (2, 3) [backward demodulation 64,112] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
415. $false <- (2, 3) [subsumption resolution 410,50] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
416. ~2 | ~3 [avatar contradiction clause 415]
430. r5_zfmisc_1(sK0,sK1,sK1,sK3,sK4,sK5,sK6) <- (2, 14) [backward demodulation 64,156] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
435. $false <- (2, 14) [subsumption resolution 430,49] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
436. ~2 | ~14 [avatar contradiction clause 435]
490. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK2,sK5,sK6) <- (2, 19) [backward demodulation 64,176] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
495. $false <- (2, 19) [subsumption resolution 490,43] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
496. ~2 | ~19 [avatar contradiction clause 495]
512. r5_zfmisc_1(sK0,sK1,sK0,sK3,sK4,sK5,sK6) <- (2, 10) [backward demodulation 64,140] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
517. $false <- (2, 10) [subsumption resolution 512,54] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
518. ~2 | ~10 [avatar contradiction clause 517]
582. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK3,sK5,sK6) <- (2, 21) [backward demodulation 64,184] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
587. $false <- (2, 21) [subsumption resolution 582,40] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
588. ~2 | ~21 [avatar contradiction clause 587]
607. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK4,sK6) <- (2, 23) [backward demodulation 64,192] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
613. $false <- (2, 23) [subsumption resolution 607,37] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
614. ~2 | ~23 [avatar contradiction clause 613]
615. sK4 = sK5 | sK3 = sK5 | sK3 = sK4 | sK2 = sK5 | sK2 = sK4 | sK2 = sK3 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1) [resolution 59,34] {a:1,w:63,wCS:6300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
616. sK4 = sK5 | sK3 = sK4 | sK2 = sK5 | sK2 = sK4 | sK2 = sK3 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~22) [subsumption resolution 615,187] {a:1,w:60,wCS:6000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
617. sK4 = sK5 | sK2 = sK5 | sK2 = sK4 | sK2 = sK3 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~21, ~22) [subsumption resolution 616,183] {a:1,w:57,wCS:5700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
618. sK4 = sK5 | sK2 = sK4 | sK2 = sK3 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~20, ~21, ~22) [subsumption resolution 617,179] {a:1,w:54,wCS:5400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
619. sK4 = sK5 | sK2 = sK3 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~19, ~20, ~21, ~22) [subsumption resolution 618,175] {a:1,w:51,wCS:5100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
620. sK4 = sK5 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 619,171] {a:1,w:48,wCS:4800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
621. sK4 = sK5 | sK1 = sK5 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 620,163] {a:1,w:45,wCS:4500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
622. sK4 = sK5 | sK1 = sK5 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 621,159] {a:1,w:42,wCS:4200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
623. sK4 = sK5 | sK1 = sK5 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 622,155] {a:1,w:39,wCS:3900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
624. sK4 = sK5 | sK1 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 623,151] {a:1,w:36,wCS:3600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
625. sK4 = sK5 | sK1 = sK5 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 624,147] {a:1,w:33,wCS:3300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
626. sK4 = sK5 | sK1 = sK5 | sK0 = sK2 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 625,143] {a:1,w:30,wCS:3000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
627. sK4 = sK5 | sK1 = sK5 | sK0 = sK1 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 626,139] {a:1,w:27,wCS:2700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
628. sK4 = sK5 | sK1 = sK5 | sK5 = sK6 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 627,135] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
629. sK4 = sK5 | sK1 = sK5 | sK4 = sK6 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 628,131] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
630. sK4 = sK5 | sK1 = sK5 | sK3 = sK6 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~7, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 629,127] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
631. sK4 = sK5 | sK1 = sK5 | sK2 = sK6 | sK1 = sK6 | sK0 = sK6 <- (~1, ~6, ~7, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 630,123] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
632. sK4 = sK5 | sK1 = sK5 | sK1 = sK6 | sK0 = sK6 <- (~1, ~5, ~6, ~7, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 631,119] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
633. sK4 = sK5 | sK1 = sK5 | sK0 = sK6 <- (~1, ~4, ~5, ~6, ~7, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 632,115] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
634. sK4 = sK5 | sK1 = sK5 <- (~1, ~3, ~4, ~5, ~6, ~7, ~8, ~9, ~10, ~11, ~12, ~13, ~14, ~15, ~16, ~18, ~19, ~20, ~21, ~22) [subsumption resolution 633,111] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
635. 17 | 23 | 1 | 3 | 4 | 5 | 6 | 7 | 8 | 9 | 10 | 11 | 12 | 13 | 14 | 15 | 16 | 18 | 19 | 20 | 21 | 22 [avatar split clause 634,186,182,178,174,170,162,158,154,150,146,142,138,134,130,126,122,118,114,110,57,190,166]
641. r5_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK1,sK6) <- (2, 17) [forward demodulation 64,168] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
642. $false <- (2, 17) [subsumption resolution 641,46] {a:0,w:0,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
643. ~2 | ~17 [avatar contradiction clause 642]
644. $false [avatar sat refutation 60,65,197,204,213,225,235,246,265,277,292,309,326,352,370,392,416,436,496,518,588,614,635,643] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:812,thDist:-812}
% SZS output end Proof for borsuk_5__t8_borsuk_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 2
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 582
% Active clauses: 128
% Passive clauses: 133
% Final active clauses: 39
% Final passive clauses: 1
% Main loop iterations started: 130
% 
% >>> Simplifying Inferences
% Duplicate literals: 105
% Fw subsumption resolutions: 280
% Fw demodulations: 2
% Bw demodulations: 69
% 
% >>> Deletion Inferences
% Equational tautologies: 21
% Forward subsumptions: 55
% 
% >>> Generating Inferences
% Binary resolution: 42
% Backward superposition: 25
% 
% >>> AVATAR
% Split clauses: 29
% Split components: 346
% Sat splitting refutations: 21
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 59
% SAT solver unit clauses: 2
% SAT solver binary clauses: 21
% 
% Memory used [KB]: 10106
% Time elapsed: 0.977 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.216 s
% congruence closure: 0.001 s
% backward subsumption index maintenance: 0.001 s
% forward subsumption: 0.003 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.968 s ( own 0.752 s ) 
% other: 0.977 s ( own 0.003 s ) 

