i: [1,49,0,3,1,1,0,0] t67_xboole_1
a: [3,69,0,3,1,49]
s: 69
i: [1,48,0,3,1,1,0,0] t67_xboole_1
a: [3,74,0,3,1,48]
s: 74
i: [1,47,0,3,1,1,0,0] t67_xboole_1
a: [3,79,0,3,1,47]
s: 79
i: [1,46,0,3,1,1,0,0] t67_xboole_1
a: [3,84,0,3,1,46]
s: 84
i: [1,50,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,89,0,2,1,50]
s: 89
i: [1,51,0,2,1,0,0,1] fc1_xboole_0
a: [3,94,0,2,1,51]
s: 94
i: [1,64,0,2,1,0,0,2] rc2_xboole_0
a: [3,104,0,2,1,64]
s: 104
i: [1,65,0,2,1,0,0,2] rc1_xboole_0
a: [3,109,0,2,1,65]
s: 109
i: [1,60,0,3,1,0,0,1] reflexivity_r1_tarski
s: 60
i: [1,52,0,3,1,0,0,255] d2_xboole_0
a: [3,99,0,3,1,52]
s: 99
r: 89
i: [1,53,0,5,1,0,0,2] t2_boole
s: 53
i: [1,54,0,5,2,0,0,1] t6_boole
s: 54
i: [1,61,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 61
i: [1,55,0,6,2,0,0,1] t3_xboole_1
s: 55
i: [1,63,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 63
i: [1,58,0,7,3,0,0,2] t8_boole
s: 58
d: [2,117,1,3,1,1,40,63,74]
a: [3,121,1,3,1,117]
s: 121
i: [1,62,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 62
d: [2,124,1,5,1,0,44,62,53]
s: 124
i: [1,56,0,8,2,0,0,1] d7_xboole_0
s: 56
i: [1,57,0,8,2,0,0,1] d7_xboole_0
s: 57
d: [2,139,1,6,2,0,44,57,53]
d: [2,146,1,3,1,0,19,139]
s: 146
i: [1,59,0,11,3,0,0,1] t19_xboole_1
s: 59
d: [2,144,2,6,2,0,44,57,124]
d: [2,145,2,3,1,0,19,144]
s: 145
d: [2,125,1,5,1,0,44,62,53]
s: 125
d: [2,126,1,5,1,0,44,53,62]
s: 126
d: [2,147,2,3,1,0,40,146,63]
s: 147
d: [2,127,1,5,1,0,44,53,62]
s: 127
d: [2,148,2,5,2,0,44,146,54]
s: 148
d: [2,133,1,6,2,0,44,56,61]
s: 133
d: [2,136,1,6,2,0,44,61,56]
s: 136
d: [2,157,3,5,2,0,44,145,54]
s: 157
d: [2,140,1,6,2,0,44,57,61]
s: 140
d: [2,158,3,5,2,0,40,148,63]
s: 158
d: [2,112,1,7,3,0,44,54,54]
s: 112
d: [2,113,1,7,2,0,44,53,54]
s: 113
d: [2,128,2,7,2,0,44,124,54]
s: 128
d: [2,116,1,8,3,0,44,55,54]
s: 116
d: [2,188,2,5,2,1,40,116,79]
a: [3,195,2,2,1,188]
s: 195
d: [2,134,1,8,2,0,44,56,62]
s: 134
d: [2,135,1,8,2,0,44,56,62]
s: 135
d: [2,189,2,5,2,1,40,116,84]
a: [3,204,2,2,1,189]
s: 204
d: [2,137,1,8,2,0,44,62,56]
s: 137
d: [2,167,2,7,2,0,44,113,62]
s: 167
d: [2,138,1,8,2,0,44,62,56]
s: 138
d: [2,142,1,8,2,0,44,57,62]
s: 142
d: [2,168,2,7,2,0,44,113,62]
s: 168
d: [2,143,1,8,2,0,44,57,62]
s: 143
d: [2,173,2,7,2,0,44,62,113]
s: 173
d: [2,151,1,12,4,0,44,59,56]
s: 151
d: [2,239,2,9,3,1,40,151,74]
s: 239
d: [2,241,3,6,2,2,40,239,79]
d: [2,243,3,3,1,3,21,241,84]
a: [3,247,3,3,1,243]
s: 247
d: [2,258,4,0,0,2,52,69,247,55]
e: 258
d: [2,259,4,0,0,3,52,94,69,247,116]
e: 259
d: [2,260,4,3,1,1,40,247,55]
d: [2,265,4,0,0,2,21,260,69]
e: 265
d: [2,261,4,5,2,1,40,247,116]
d: [2,267,4,2,1,2,21,261,69]
d: [2,268,4,0,0,3,21,267,94]
e: 268
f: 265,243,46,47,48,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t67_xboole_1
% SZS output start Proof for xboole_1__t67_xboole_1
1. ! [X0] : ! [X1] : ! [X2] : ((r1_xboole_0(X1,X2) & r1_tarski(X0,X2) & r1_tarski(X0,X1)) => k1_xboole_0 = X0) [input t67_xboole_1]
2. ~! [X0] : ! [X1] : ! [X2] : ((r1_xboole_0(X1,X2) & r1_tarski(X0,X2) & r1_tarski(X0,X1)) => k1_xboole_0 = X0) [negated conjecture 1]
6. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) <=> k1_xboole_0 = k3_xboole_0(X0,X1)) [input d7_xboole_0]
16. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k3_xboole_0(X1,X2))) [input t19_xboole_1]
18. ! [X0] : (r1_tarski(X0,k1_xboole_0) => k1_xboole_0 = X0) [input t3_xboole_1]
22. ~! [X0,X1,X2] : ((r1_xboole_0(X1,X2) & r1_tarski(X0,X2) & r1_tarski(X0,X1)) => k1_xboole_0 = X0) [flattening 2]
23. ! [X0,X1] : (r1_xboole_0(X0,X1) <=> k1_xboole_0 = k3_xboole_0(X0,X1)) [flattening 6]
26. ! [X0,X1,X2] : ((r1_tarski(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k3_xboole_0(X1,X2))) [flattening 16]
31. ? [X0,X1,X2] : (k1_xboole_0 != X0 & (r1_xboole_0(X1,X2) & r1_tarski(X0,X2) & r1_tarski(X0,X1))) [ennf transformation 22]
32. ? [X0,X1,X2] : (k1_xboole_0 != X0 & r1_xboole_0(X1,X2) & r1_tarski(X0,X2) & r1_tarski(X0,X1)) [flattening 31]
34. ! [X0] : (k1_xboole_0 = X0 | ~r1_tarski(X0,k1_xboole_0)) [ennf transformation 18]
36. ! [X0,X1,X2] : (r1_tarski(X0,k3_xboole_0(X1,X2)) | (~r1_tarski(X0,X2) | ~r1_tarski(X0,X1))) [ennf transformation 26]
37. ! [X0,X1,X2] : (r1_tarski(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X2) | ~r1_tarski(X0,X1)) [flattening 36]
39. ? [X0,X1,X2] : (k1_xboole_0 != X0 & r1_xboole_0(X1,X2) & r1_tarski(X0,X2) & r1_tarski(X0,X1)) => (k1_xboole_0 != sK0 & r1_xboole_0(sK1,sK2) & r1_tarski(sK0,sK2) & r1_tarski(sK0,sK1)) [choice axiom]
40. k1_xboole_0 != sK0 & r1_xboole_0(sK1,sK2) & r1_tarski(sK0,sK2) & r1_tarski(sK0,sK1) [skolemisation 32,39]
41. ! [X0,X1] : ((r1_xboole_0(X0,X1) | k1_xboole_0 != k3_xboole_0(X0,X1)) & (k1_xboole_0 = k3_xboole_0(X0,X1) | ~r1_xboole_0(X0,X1))) [nnf transformation 23]
46. r1_tarski(sK0,sK1) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. r1_tarski(sK0,sK2) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. r1_xboole_0(sK1,sK2) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. k1_xboole_0 != sK0 [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. ~r1_tarski(X0,k1_xboole_0) | k1_xboole_0 = X0 [cnf transformation 34] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. k1_xboole_0 = k3_xboole_0(X0,X1) | ~r1_xboole_0(X0,X1) [cnf transformation 41] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. r1_tarski(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 37] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. 1 <=> k1_xboole_0 = sK0 [avatar definition]
69. k1_xboole_0 != sK0 <- (~1) [avatar component clause 67] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
70. ~1 [avatar split clause 49,67]
72. 2 <=> r1_xboole_0(sK1,sK2) [avatar definition]
74. r1_xboole_0(sK1,sK2) <- (2) [avatar component clause 72] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. 2 [avatar split clause 48,72]
77. 3 <=> r1_tarski(sK0,sK2) [avatar definition]
79. r1_tarski(sK0,sK2) <- (3) [avatar component clause 77] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. 3 [avatar split clause 47,77]
82. 4 <=> r1_tarski(sK0,sK1) [avatar definition]
84. r1_tarski(sK0,sK1) <- (4) [avatar component clause 82] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
85. 4 [avatar split clause 46,82]
151. ~r1_xboole_0(X4,X5) | ~r1_tarski(X6,X5) | ~r1_tarski(X6,X4) | r1_tarski(X6,k1_xboole_0) [superposition 59,56] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
239. ~r1_tarski(X10,sK2) | ~r1_tarski(X10,sK1) | r1_tarski(X10,k1_xboole_0) <- (2) [resolution 151,74] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
241. ~r1_tarski(sK0,sK1) | r1_tarski(sK0,k1_xboole_0) <- (2, 3) [resolution 239,79] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
243. r1_tarski(sK0,k1_xboole_0) <- (2, 3, 4) [subsumption resolution 241,84] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
245. 15 <=> r1_tarski(sK0,k1_xboole_0) [avatar definition]
247. r1_tarski(sK0,k1_xboole_0) <- (15) [avatar component clause 245] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
248. 15 | ~2 | ~3 | ~4 [avatar split clause 243,82,77,72,245]
260. k1_xboole_0 = sK0 <- (15) [resolution 247,55] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
265. $false <- (~1, 15) [subsumption resolution 260,69] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
266. 1 | ~15 [avatar contradiction clause 265]
270. $false [avatar sat refutation 70,75,80,85,248,266] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:96,thDist:-96}
% SZS output end Proof for xboole_1__t67_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 20
% Generated clauses: 152
% Active clauses: 38
% Passive clauses: 60
% Final active clauses: 37
% Final passive clauses: 8
% Main loop iterations started: 52
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Trivial inequalities: 6
% Fw subsumption resolutions: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 4
% Forward subsumptions: 62
% 
% >>> Generating Inferences
% Binary resolution: 25
% Unit resulting resolution: 2
% Forward superposition: 48
% Backward superposition: 29
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 9
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

