i: [1,28,0,2,1,1,0,0] t26_xxreal_3
a: [3,50,0,2,1,28]
s: 50
i: [1,27,0,2,1,1,0,0] t26_xxreal_3
a: [3,55,0,2,1,27]
s: 55
i: [1,37,0,2,1,0,0,1] rc1_xxreal_0
a: [3,60,0,2,1,37]
s: 60
i: [1,29,0,17,2,1,0,0] t26_xxreal_3
a: [3,41,0,9,1,29]
s: 41
i: [1,30,0,5,2,0,0,1] dt_k2_xxreal_3
s: 30
i: [1,31,0,7,2,0,0,1] involutiveness_k2_xxreal_3
s: 31
i: [1,34,0,8,3,0,0,1] dt_k3_xxreal_3
s: 34
i: [1,35,0,8,3,0,0,1] dt_k1_xxreal_3
s: 35
i: [1,36,0,11,3,0,0,1] commutativity_k1_xxreal_3
s: 36
i: [1,32,0,12,3,0,0,1] d4_xxreal_3
s: 32
i: [1,33,0,13,3,0,0,1] t25_xxreal_3
s: 33
d: [2,63,1,14,3,1,44,41,36]
d: [2,69,1,12,2,2,21,63,50]
a: [3,73,1,3,1,69]
s: 73
d: [2,102,2,0,0,2,52,55,73,30]
e: 102
d: [2,103,2,2,1,1,40,73,30]
d: [2,105,2,0,0,2,21,103,55]
e: 105
r: 73
a: [3,72,1,3,1,69]
s: 72
a: [3,77,1,9,1,69]
s: 77
a: [3,45,0,8,1,29]
s: 45
d: [2,98,1,14,5,0,44,34,33]
d: [2,99,1,12,4,0,18,98]
d: [2,101,1,9,3,0,21,99,30]
s: 101
d: [2,113,2,14,5,0,44,101,32]
d: [2,114,2,12,4,0,18,113]
d: [2,117,2,9,3,0,21,114,30]
s: 117
d: [2,81,1,15,4,0,44,32,31]
d: [2,92,1,12,3,0,21,81,30]
s: 92
d: [2,118,3,15,5,0,44,117,33]
d: [2,119,3,13,4,0,18,118]
d: [2,120,3,10,3,0,21,119,30]
s: 120
d: [2,82,1,17,5,0,44,32,36]
d: [2,91,1,15,4,0,18,82]
d: [2,93,1,12,3,0,21,91,30]
s: 93
d: [2,83,1,17,5,0,44,32,36]
d: [2,90,1,15,4,0,18,83]
d: [2,94,1,12,3,0,21,90,30]
s: 94
d: [2,134,4,15,5,0,44,120,32]
d: [2,136,4,13,4,0,18,134]
d: [2,139,4,10,3,0,21,136,30]
s: 139
d: [2,84,1,17,5,0,44,36,32]
d: [2,89,1,15,4,0,18,84]
d: [2,95,1,12,3,0,21,89,30]
s: 95
d: [2,145,4,15,5,0,44,120,93]
d: [2,155,4,13,4,0,18,145]
d: [2,161,4,10,3,0,21,155,30]
s: 161
d: [2,85,1,17,5,0,44,36,32]
d: [2,88,1,15,4,0,18,85]
d: [2,96,1,12,3,0,21,88,30]
s: 96
d: [2,97,1,16,4,0,44,33,31]
d: [2,100,1,13,3,0,21,97,30]
s: 100
d: [2,164,5,16,5,0,44,139,33]
d: [2,168,5,14,4,0,18,164]
d: [2,169,5,11,3,0,21,168,30]
s: 169
d: [2,140,2,15,4,0,44,93,31]
d: [2,158,2,12,3,0,21,140,30]
s: 158
d: [2,195,6,16,5,0,44,169,32]
d: [2,202,6,14,4,0,18,195]
d: [2,205,6,11,3,0,21,202,30]
s: 205
d: [2,172,2,16,5,0,44,100,32]
d: [2,188,2,12,3,0,18,172]
s: 188
d: [2,234,3,0,0,3,52,50,55,45,188]
e: 234
d: [2,238,3,11,3,1,44,45,188]
d: [2,253,3,4,2,1,19,238]
d: [2,260,3,2,1,2,21,253,55]
d: [2,261,3,0,0,3,21,260,50]
e: 261
d: [2,65,1,14,3,1,44,41,36]
d: [2,79,1,12,2,2,21,65,50]
d: [2,107,2,12,3,1,44,77,32]
d: [2,108,2,10,2,2,21,107,55]
d: [2,109,2,8,1,3,21,108,50]
f: 261,109,105,79,27,28,29
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_3__t26_xxreal_3
% SZS output start Proof for xxreal_3__t26_xxreal_3
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => (k2_xxreal_3(k3_xxreal_3(X0,X1)) = k3_xxreal_3(X1,X0) & k2_xxreal_3(k3_xxreal_3(X0,X1)) = k1_xxreal_3(k2_xxreal_3(X0),X1)))) [input t26_xxreal_3]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => (k2_xxreal_3(k3_xxreal_3(X0,X1)) = k3_xxreal_3(X1,X0) & k2_xxreal_3(k3_xxreal_3(X0,X1)) = k1_xxreal_3(k2_xxreal_3(X0),X1)))) [negated conjecture 1]
3. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => k1_xxreal_3(X0,X1) = k1_xxreal_3(X1,X0)) [input commutativity_k1_xxreal_3]
4. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => k3_xxreal_3(X0,X1) = k1_xxreal_3(X0,k2_xxreal_3(X1)))) [input d4_xxreal_3]
6. ! [X0] : (v1_xxreal_0(X0) => v1_xxreal_0(k2_xxreal_3(X0))) [input dt_k2_xxreal_3]
8. ! [X0] : (v1_xxreal_0(X0) => k2_xxreal_3(k2_xxreal_3(X0)) = X0) [input involutiveness_k2_xxreal_3]
10. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => k2_xxreal_3(k1_xxreal_3(X0,X1)) = k3_xxreal_3(k2_xxreal_3(X1),X0))) [input t25_xxreal_3]
11. ? [X0] : (? [X1] : ((k2_xxreal_3(k3_xxreal_3(X0,X1)) != k3_xxreal_3(X1,X0) | k2_xxreal_3(k3_xxreal_3(X0,X1)) != k1_xxreal_3(k2_xxreal_3(X0),X1)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
12. ! [X0] : (v1_xxreal_0(k2_xxreal_3(X0)) | ~v1_xxreal_0(X0)) [ennf transformation 6]
13. ! [X0] : (k2_xxreal_3(k2_xxreal_3(X0)) = X0 | ~v1_xxreal_0(X0)) [ennf transformation 8]
14. ! [X0] : (! [X1] : (k3_xxreal_3(X0,X1) = k1_xxreal_3(X0,k2_xxreal_3(X1)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 4]
15. ! [X0] : (! [X1] : (k2_xxreal_3(k1_xxreal_3(X0,X1)) = k3_xxreal_3(k2_xxreal_3(X1),X0) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 10]
20. ! [X0,X1] : (k1_xxreal_3(X0,X1) = k1_xxreal_3(X1,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 3]
21. ! [X0,X1] : (k1_xxreal_3(X0,X1) = k1_xxreal_3(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 20]
22. ? [X0] : (? [X1] : ((k2_xxreal_3(k3_xxreal_3(X0,X1)) != k3_xxreal_3(X1,X0) | k2_xxreal_3(k3_xxreal_3(X0,X1)) != k1_xxreal_3(k2_xxreal_3(X0),X1)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : ((k2_xxreal_3(k3_xxreal_3(sK0,X1)) != k3_xxreal_3(X1,sK0) | k2_xxreal_3(k3_xxreal_3(sK0,X1)) != k1_xxreal_3(k2_xxreal_3(sK0),X1)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
23. ? [X1] : ((k2_xxreal_3(k3_xxreal_3(sK0,X1)) != k3_xxreal_3(X1,sK0) | k2_xxreal_3(k3_xxreal_3(sK0,X1)) != k1_xxreal_3(k2_xxreal_3(sK0),X1)) & v1_xxreal_0(X1)) => ((k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k3_xxreal_3(sK1,sK0) | k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k1_xxreal_3(k2_xxreal_3(sK0),sK1)) & v1_xxreal_0(sK1)) [choice axiom]
24. ((k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k3_xxreal_3(sK1,sK0) | k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k1_xxreal_3(k2_xxreal_3(sK0),sK1)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 11,23,22]
27. v1_xxreal_0(sK0) [cnf transformation 24] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
28. v1_xxreal_0(sK1) [cnf transformation 24] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
29. k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k3_xxreal_3(sK1,sK0) | k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k1_xxreal_3(k2_xxreal_3(sK0),sK1) [cnf transformation 24] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
30. v1_xxreal_0(k2_xxreal_3(X0)) | ~v1_xxreal_0(X0) [cnf transformation 12] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
31. k2_xxreal_3(k2_xxreal_3(X0)) = X0 | ~v1_xxreal_0(X0) [cnf transformation 13] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
32. k3_xxreal_3(X0,X1) = k1_xxreal_3(X0,k2_xxreal_3(X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 14] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. k2_xxreal_3(k1_xxreal_3(X0,X1)) = k3_xxreal_3(k2_xxreal_3(X1),X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 15] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
36. k1_xxreal_3(X0,X1) = k1_xxreal_3(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 21] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
39. 1 <=> k2_xxreal_3(k3_xxreal_3(sK0,sK1)) = k1_xxreal_3(k2_xxreal_3(sK0),sK1) [avatar definition]
41. k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k1_xxreal_3(k2_xxreal_3(sK0),sK1) <- (~1) [avatar component clause 39] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
43. 2 <=> k2_xxreal_3(k3_xxreal_3(sK0,sK1)) = k3_xxreal_3(sK1,sK0) [avatar definition]
45. k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k3_xxreal_3(sK1,sK0) <- (~2) [avatar component clause 43] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
46. ~1 | ~2 [avatar split clause 29,43,39]
48. 3 <=> v1_xxreal_0(sK1) [avatar definition]
50. v1_xxreal_0(sK1) <- (3) [avatar component clause 48] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
51. 3 [avatar split clause 28,48]
53. 4 <=> v1_xxreal_0(sK0) [avatar definition]
55. v1_xxreal_0(sK0) <- (4) [avatar component clause 53] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
56. 4 [avatar split clause 27,53]
65. k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k1_xxreal_3(sK1,k2_xxreal_3(sK0)) | ~v1_xxreal_0(k2_xxreal_3(sK0)) | ~v1_xxreal_0(sK1) <- (~1) [superposition 41,36] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
71. 6 <=> v1_xxreal_0(k2_xxreal_3(sK0)) [avatar definition]
73. ~v1_xxreal_0(k2_xxreal_3(sK0)) <- (~6) [avatar component clause 71] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
75. 7 <=> k2_xxreal_3(k3_xxreal_3(sK0,sK1)) = k1_xxreal_3(sK1,k2_xxreal_3(sK0)) [avatar definition]
77. k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k1_xxreal_3(sK1,k2_xxreal_3(sK0)) <- (~7) [avatar component clause 75] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
79. k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k1_xxreal_3(sK1,k2_xxreal_3(sK0)) | ~v1_xxreal_0(k2_xxreal_3(sK0)) <- (~1, 3) [subsumption resolution 65,50] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
80. ~6 | ~7 | 1 | ~3 [avatar split clause 79,48,39,75,71]
97. k3_xxreal_3(X0,X1) = k2_xxreal_3(k1_xxreal_3(X1,k2_xxreal_3(X0))) | ~v1_xxreal_0(k2_xxreal_3(X0)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [superposition 33,31] {a:1,w:16,wCS:6400,sine:1,thAx:0,allAx:2,thDist:-2}
100. k3_xxreal_3(X0,X1) = k2_xxreal_3(k1_xxreal_3(X1,k2_xxreal_3(X0))) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [subsumption resolution 97,30] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
103. ~v1_xxreal_0(sK0) <- (~6) [resolution 73,30] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
105. $false <- (4, ~6) [subsumption resolution 103,55] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
106. ~4 | 6 [avatar contradiction clause 105]
107. k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k3_xxreal_3(sK1,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~7) [superposition 77,32] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
108. k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k3_xxreal_3(sK1,sK0) | ~v1_xxreal_0(sK1) <- (4, ~7) [subsumption resolution 107,55] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
109. k2_xxreal_3(k3_xxreal_3(sK0,sK1)) != k3_xxreal_3(sK1,sK0) <- (3, 4, ~7) [subsumption resolution 108,50] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
110. ~2 | ~3 | ~4 | 7 [avatar split clause 109,75,53,48,43]
172. k2_xxreal_3(k3_xxreal_3(X0,X1)) = k3_xxreal_3(X1,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [superposition 100,32] {a:2,w:16,wCS:6400,sine:1,thAx:0,allAx:3,thDist:-3}
188. k2_xxreal_3(k3_xxreal_3(X0,X1)) = k3_xxreal_3(X1,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) [duplicate literal removal 172] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
238. k3_xxreal_3(sK1,sK0) != k3_xxreal_3(sK1,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~2) [superposition 45,188] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
253. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~2) [trivial inequality removal 238] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
260. ~v1_xxreal_0(sK1) <- (~2, 4) [subsumption resolution 253,55] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
261. $false <- (~2, 3, 4) [subsumption resolution 260,50] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
262. 2 | ~3 | ~4 [avatar contradiction clause 261]
269. $false [avatar sat refutation 46,51,56,80,106,110,262] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:72,thDist:-72}
% SZS output end Proof for xxreal_3__t26_xxreal_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 11
% Generated clauses: 211
% Active clauses: 26
% Passive clauses: 48
% Final active clauses: 25
% Final passive clauses: 18
% Main loop iterations started: 30
% 
% >>> Simplifying Inferences
% Duplicate literals: 64
% Trivial inequalities: 1
% Fw subsumption resolutions: 47
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 45
% 
% >>> Generating Inferences
% Binary resolution: 1
% Unit resulting resolution: 2
% Forward superposition: 46
% Backward superposition: 44
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 3
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% splitting model update: 0.001 s ( own 0.0000 s ) 
% backward superposition index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

