i: [1,85,0,2,1,1,0,0] t48_funct_4
a: [3,140,0,2,1,85]
s: 140
i: [1,87,0,12,2,1,0,0] t48_funct_4
a: [3,126,0,3,1,87]
s: 126
i: [1,86,0,8,1,1,0,0] t48_funct_4
a: [3,135,0,8,1,86]
s: 135
i: [1,122,0,2,1,0,0,1] rc1_funct_1
a: [3,145,0,2,1,122]
s: 145
i: [1,121,0,2,1,0,0,1] rc1_funct_1
a: [3,150,0,2,1,121]
s: 150
i: [1,98,0,3,1,0,0,2] reflexivity_r1_tarski
s: 98
i: [1,97,0,4,1,0,0,1] existence_m1_subset_1
s: 97
i: [1,99,0,4,1,0,0,1] fc6_relat_1
s: 99
i: [1,109,0,4,1,0,0,255] rc3_relat_1
s: 109
i: [1,112,0,4,1,0,0,255] rc6_funct_1
s: 112
i: [1,115,0,4,1,0,0,255] rc6_funct_1
s: 115
i: [1,110,0,5,1,0,0,255] rc3_relat_1
s: 110
i: [1,111,0,5,1,0,0,255] rc3_relat_1
s: 111
i: [1,113,0,5,1,0,0,255] rc6_funct_1
s: 113
i: [1,114,0,5,1,0,0,255] rc6_funct_1
s: 114
i: [1,89,0,7,3,0,0,1] dt_k2_funct_4
s: 89
i: [1,90,0,7,3,0,0,1] dt_k2_funct_4
s: 90
i: [1,92,0,7,2,0,0,2] t3_subset
s: 92
d: [2,152,1,4,2,1,40,90,126]
d: [2,153,1,2,1,2,21,152,140]
a: [3,157,1,2,1,153]
s: 157
i: [1,93,0,7,2,0,0,2] t3_subset
s: 93
i: [1,88,0,8,3,0,0,1] cc2_relat_1
s: 88
d: [2,167,2,0,0,2,52,157,99,135,88]
e: 167
d: [2,168,1,8,2,1,40,88,135]
d: [2,172,1,6,1,2,21,168,157]
d: [2,173,1,0,0,2,21,172,99]
e: 173
r: 126
r: 157
a: [3,156,1,2,1,153]
s: 156
a: [3,125,0,3,1,87]
s: 125
d: [2,159,1,7,1,1,40,92,135]
a: [3,164,1,7,1,159]
s: 164
a: [3,130,0,9,1,87]
s: 130
i: [1,106,0,8,2,0,0,1] cc1_relset_1
s: 106
d: [2,177,1,8,1,1,40,130,93]
a: [3,181,1,8,1,177]
s: 181
i: [1,107,0,9,2,0,0,255] cc2_relset_1
s: 107
d: [2,186,1,5,1,1,40,107,135]
a: [3,192,1,5,1,186]
s: 192
i: [1,108,0,9,2,0,0,255] cc2_relset_1
s: 108
i: [1,91,0,10,4,0,0,1] cc3_funct_1
s: 91
d: [2,194,1,3,1,1,40,108,135]
a: [3,200,1,3,1,194]
s: 200
i: [1,100,0,11,3,0,0,255] redefinition_k2_relset_1
s: 100
d: [2,160,1,5,1,0,40,92,97]
s: 160
i: [1,101,0,11,3,0,0,255] redefinition_k1_relset_1
s: 101
i: [1,102,0,11,3,0,0,255] dt_k1_relset_1
s: 102
d: [2,169,1,6,2,0,40,88,97]
s: 169
i: [1,103,0,11,3,0,0,255] dt_k2_relset_1
s: 103
d: [2,184,1,6,1,0,40,106,97]
s: 184
i: [1,116,0,11,2,0,0,255] reflexivity_r1_relset_1
s: 116
i: [1,119,0,11,2,0,0,255] fc9_relset_1
s: 119
d: [2,224,1,7,1,1,40,116,135]
a: [3,229,1,7,1,224]
s: 229
i: [1,120,0,11,2,0,0,2] fc8_relset_1
s: 120
d: [2,238,1,3,1,1,40,120,135]
a: [3,243,1,3,1,238]
s: 243
i: [1,104,0,12,4,0,0,255] cc5_relat_1
s: 104
i: [1,105,0,12,4,0,0,255] cc6_relat_1
s: 105
d: [2,248,1,16,3,1,40,104,135]
d: [2,250,1,10,2,1,21,248,99]
s: 250
i: [1,117,0,14,3,0,0,255] redefinition_r1_relset_1
s: 117
d: [2,254,1,16,3,1,40,105,135]
d: [2,256,1,10,2,1,21,254,99]
s: 256
i: [1,118,0,14,3,0,0,255] redefinition_r1_relset_1
s: 118
i: [1,94,0,16,4,0,0,255] t4_relset_1
s: 94
d: [2,280,1,15,3,1,40,94,130]
a: [3,287,1,3,1,280]
s: 287
d: [2,297,2,0,0,3,52,156,140,287,89]
e: 297
d: [2,298,2,4,2,1,40,287,89]
d: [2,300,2,2,1,2,21,298,140]
d: [2,301,2,0,0,3,21,300,156]
e: 301
r: 287
i: [1,95,0,16,4,0,0,255] t47_funct_4
s: 95
a: [3,286,1,3,1,280]
s: 286
i: [1,96,0,17,4,0,0,1] t45_funct_4
s: 96
a: [3,291,1,7,1,280]
s: 291
d: [2,306,2,10,3,1,40,291,96]
d: [2,307,2,8,2,2,21,306,140]
d: [2,308,2,6,1,3,21,307,156]
a: [3,312,2,6,1,308]
s: 312
d: [2,260,1,10,2,1,40,117,135]
s: 260
d: [2,265,1,10,2,1,40,118,135]
s: 265
d: [2,170,1,7,3,0,40,88,93]
s: 170
d: [2,185,1,7,2,0,40,106,93]
s: 185
d: [2,187,1,7,1,0,40,107,97]
s: 187
d: [2,195,1,7,1,0,40,108,97]
s: 195
d: [2,188,1,8,2,0,40,107,93]
s: 188
d: [2,196,1,8,2,0,40,108,93]
s: 196
d: [2,203,1,8,3,0,40,91,97]
s: 203
d: [2,325,2,5,1,0,40,188,98]
s: 325
d: [2,204,1,9,4,0,40,91,93]
s: 204
d: [2,329,2,5,1,0,40,196,98]
s: 329
d: [2,234,1,9,1,0,40,119,97]
s: 234
d: [2,239,1,9,1,0,40,120,97]
s: 239
d: [2,207,1,10,3,0,40,102,92]
s: 207
d: [2,211,1,15,5,0,44,102,101]
d: [2,212,1,10,3,0,18,211]
s: 212
d: [2,345,2,14,5,0,44,207,101]
d: [2,346,2,9,3,0,18,345]
s: 346
d: [2,360,3,0,0,3,52,156,192,312,346]
e: 360
d: [2,363,3,7,2,1,40,346,312]
d: [2,372,3,5,1,2,21,363,156]
d: [2,373,3,0,0,3,21,372,192]
e: 373
r: 291
r: 312
a: [3,295,1,5,1,280]
s: 295
a: [3,290,1,7,1,280]
s: 290
d: [2,378,2,4,1,1,40,290,185]
a: [3,401,2,4,1,378]
s: 401
d: [2,215,1,10,3,0,40,103,92]
s: 215
d: [2,379,2,12,3,1,40,290,204]
d: [2,403,2,8,2,1,21,379,99]
a: [3,411,2,4,1,403]
s: 411
d: [2,219,1,15,5,0,44,103,100]
d: [2,220,1,10,3,0,18,219]
s: 220
d: [2,221,1,10,2,0,40,116,93]
s: 221
d: [2,376,2,5,1,1,40,290,196]
a: [3,391,2,5,1,376]
s: 391
d: [2,231,1,10,2,0,40,119,93]
s: 231
d: [2,377,2,5,1,1,40,290,188]
a: [3,396,2,5,1,377]
s: 396
d: [2,235,1,10,2,0,40,120,93]
s: 235
d: [2,249,1,10,3,0,40,104,97]
s: 249
a: [3,311,2,6,1,308]
s: 311
d: [2,255,1,10,3,0,40,105,97]
s: 255
d: [2,463,3,4,1,1,40,311,196]
a: [3,483,3,4,1,463]
s: 483
d: [2,206,1,11,4,0,40,102,88]
s: 206
d: [2,214,1,11,4,0,40,103,88]
s: 214
d: [2,464,3,4,1,1,40,311,188]
a: [3,488,3,4,1,464]
s: 488
d: [2,245,1,11,4,0,40,104,93]
s: 245
d: [2,466,3,11,3,1,40,311,204]
d: [2,490,3,7,2,1,21,466,99]
a: [3,498,3,4,1,490]
s: 498
d: [2,251,1,11,4,0,40,105,93]
s: 251
d: [2,205,1,13,5,0,40,102,91]
s: 205
d: [2,461,3,10,3,1,40,311,95]
d: [2,468,3,8,2,2,21,461,140]
d: [2,469,3,6,1,3,21,468,156]
a: [3,473,3,6,1,469]
s: 473
r: 295
d: [2,530,1,4,1,2,23,295,473]
a: [3,535,1,4,1,530]
s: 535
d: [2,462,3,7,1,1,40,311,221]
a: [3,478,3,7,1,462]
s: 478
d: [2,210,1,13,3,0,40,102,106]
s: 210
d: [2,213,1,13,5,0,40,103,91]
s: 213
d: [2,375,2,14,3,1,40,290,95]
d: [2,381,2,11,2,2,21,375,125]
d: [2,382,2,8,1,3,21,381,286]
a: [3,386,2,8,1,382]
s: 386
d: [2,218,1,13,3,0,40,103,106]
s: 218
d: [2,542,2,7,1,2,22,386,473]
a: [3,546,2,7,1,542]
s: 546
d: [2,225,1,13,1,0,40,116,97]
s: 225
d: [2,257,1,13,3,0,40,117,93]
s: 257
d: [2,550,3,14,3,1,44,220,546]
a: [3,554,3,4,1,550]
s: 554
d: [2,568,4,0,0,3,52,286,125,554,89]
e: 568
d: [2,569,4,6,2,1,40,554,89]
d: [2,571,4,3,1,2,21,569,125]
d: [2,572,4,0,0,3,21,571,286]
e: 572
r: 554
d: [2,262,1,13,3,0,40,118,93]
s: 262
a: [3,553,3,4,1,550]
s: 553
d: [2,270,1,13,4,0,40,94,108]
s: 270
d: [2,271,1,13,4,0,40,94,107]
s: 271
d: [2,603,4,16,4,1,44,271,473]
d: [2,605,4,13,3,2,21,603,286]
a: [3,611,4,4,1,605]
s: 611
r: 535
d: [2,619,5,0,0,1,40,611,98]
e: 619
d: [2,618,5,0,0,1,52,98,611]
e: 618
r: 611
s: 535
d: [2,584,4,16,4,1,44,270,473]
d: [2,586,4,13,3,2,21,584,286]
a: [3,589,4,5,1,586]
s: 589
d: [2,627,5,0,0,1,40,589,98]
e: 627
d: [2,625,5,0,0,2,40,589,290]
e: 625
d: [2,622,5,0,0,2,52,290,589]
e: 622
d: [2,623,5,0,0,4,52,156,140,311,589,96]
e: 623
r: 589
d: [2,208,1,14,3,0,40,102,108]
s: 208
d: [2,209,1,14,3,0,40,102,107]
s: 209
d: [2,585,3,19,4,1,44,270,546]
d: [2,594,3,15,3,2,21,585,553]
a: [3,597,3,6,1,594]
s: 597
d: [2,643,4,0,0,1,40,597,98]
e: 643
d: [2,639,4,0,0,1,52,98,597]
e: 639
d: [2,640,4,0,0,4,52,286,125,290,597,96]
e: 640
r: 597
d: [2,216,1,14,3,0,40,103,108]
s: 216
a: [3,592,4,8,2,586]
s: 592
d: [2,217,1,14,3,0,40,103,107]
s: 217
d: [2,246,1,15,5,0,40,104,102]
s: 246
d: [2,652,5,5,1,1,40,592,98]
a: [3,656,5,5,1,652]
s: 656
d: [2,247,1,15,5,0,40,104,103]
s: 247
d: [2,439,2,9,1,1,40,221,290]
a: [3,446,2,9,1,439]
s: 446
d: [2,252,1,15,5,0,40,105,102]
s: 252
d: [2,253,1,15,5,0,40,105,103]
s: 253
d: [2,509,3,13,3,1,40,245,311]
d: [2,515,3,9,2,1,21,509,99]
s: 515
d: [2,269,1,15,4,0,40,94,116]
s: 269
d: [2,521,3,13,3,1,40,251,311]
d: [2,527,3,9,2,1,21,521,99]
s: 527
d: [2,273,1,15,4,0,40,94,119]
s: 273
d: [2,274,1,15,4,0,40,94,120]
s: 274
a: [3,600,3,9,2,594]
s: 600
d: [2,279,1,15,4,0,40,94,92]
s: 279
d: [2,681,5,19,4,1,40,279,592]
a: [3,694,5,3,1,681]
s: 694
d: [2,261,1,16,2,0,40,117,97]
s: 261
d: [2,266,1,16,2,0,40,118,97]
s: 266
d: [2,531,4,12,3,1,44,220,473]
d: [2,537,4,9,2,2,21,531,286]
s: 537
d: [2,277,1,20,6,0,40,94,91]
d: [2,283,1,16,5,0,21,277,99]
s: 283
d: [2,707,5,11,4,2,40,537,91]
a: [3,721,5,3,1,707]
s: 721
d: [2,232,1,18,3,0,40,119,102]
s: 232
d: [2,233,1,18,3,0,40,119,103]
s: 233
d: [2,747,3,22,5,1,44,283,546]
d: [2,748,3,18,4,2,21,747,553]
a: [3,752,3,4,1,748]
s: 752
d: [2,236,1,18,3,0,40,120,102]
s: 236
d: [2,716,5,12,2,2,40,537,119]
a: [3,731,5,4,1,716]
s: 731
d: [2,237,1,18,3,0,40,120,103]
s: 237
d: [2,267,1,18,5,0,40,94,118]
s: 267
d: [2,717,5,12,2,2,40,537,120]
a: [3,739,5,4,1,717]
s: 739
d: [2,268,1,18,5,0,40,94,117]
s: 268
d: [2,708,5,9,3,2,40,537,88]
d: [2,726,5,6,2,3,21,708,694]
s: 726
d: [2,275,1,22,6,0,40,94,105]
d: [2,281,1,18,5,0,21,275,99]
s: 281
d: [2,276,1,22,6,0,40,94,104]
d: [2,282,1,18,5,0,21,276,99]
s: 282
d: [2,715,5,9,2,2,40,537,106]
d: [2,727,5,6,1,3,21,715,694]
s: 727
d: [2,303,1,24,6,0,40,96,95]
d: [2,304,1,21,5,0,21,303,90]
d: [2,305,1,18,4,0,21,304,89]
s: 305
d: [2,709,5,8,2,2,40,537,92]
s: 709
d: [2,222,1,20,3,0,40,116,102]
s: 222
d: [2,223,1,20,3,0,40,116,103]
s: 223
a: [3,608,4,9,2,605]
s: 608
d: [2,258,1,23,4,0,40,117,102]
s: 258
d: [2,794,5,6,1,2,40,608,290]
a: [3,803,5,6,1,794]
s: 803
d: [2,259,1,23,4,0,40,117,103]
s: 259
d: [2,263,1,23,4,0,40,118,102]
s: 263
d: [2,796,5,6,1,1,40,608,98]
a: [3,808,5,6,1,796]
s: 808
d: [2,264,1,23,4,0,40,118,103]
s: 264
d: [2,508,2,14,3,1,40,245,290]
d: [2,514,2,10,2,1,21,508,99]
s: 514
d: [2,520,2,14,3,1,40,251,290]
d: [2,526,2,10,2,1,21,520,99]
s: 526
d: [2,563,2,12,2,1,40,257,290]
s: 563
a: [3,557,3,10,2,550]
s: 557
d: [2,578,2,12,2,1,40,262,290]
s: 578
d: [2,564,3,10,2,1,40,257,311]
s: 564
d: [2,448,2,7,1,0,40,231,98]
s: 448
d: [2,454,2,7,1,0,40,235,98]
s: 454
d: [2,579,3,10,2,1,40,262,311]
s: 579
d: [2,351,2,9,3,0,40,212,92]
s: 351
d: [2,713,5,10,2,2,40,537,108]
s: 713
d: [2,420,2,14,5,0,44,215,100]
d: [2,421,2,9,3,0,18,420]
s: 421
d: [2,846,3,0,0,3,52,156,200,535,421]
e: 846
d: [2,863,3,5,2,1,40,421,535]
d: [2,872,3,3,1,2,21,863,156]
d: [2,873,3,0,0,3,21,872,200]
e: 873
d: [2,175,1,2,1,1,21,168,99]
f: 873,530,469,373,308,301,280,194,186,175,153,85,86,87
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t48_funct_4
% SZS output start Proof for funct_4__t48_funct_4
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2))) & v1_funct_1(X3)) => (m1_subset_1(k2_funct_4(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X0),X2))) & v1_funct_1(k2_funct_4(X3)))) [input t48_funct_4]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2))) & v1_funct_1(X3)) => (m1_subset_1(k2_funct_4(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X0),X2))) & v1_funct_1(k2_funct_4(X3)))) [negated conjecture 1]
4. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
5. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [input cc2_relset_1]
10. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => m1_subset_1(k1_relset_1(X0,X1),k1_zfmisc_1(X0))) [input dt_k1_relset_1]
12. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k2_funct_4(X0)) & v1_relat_1(k2_funct_4(X0)))) [input dt_k2_funct_4]
13. ! [X0,X1] : ((v5_relat_1(X1,X0) & v1_relat_1(X1)) => m1_subset_1(k2_relset_1(X0,X1),k1_zfmisc_1(X0))) [input dt_k2_relset_1]
18. ! [X0,X1] : v1_relat_1(k2_zfmisc_1(X0,X1)) [input fc6_relat_1]
24. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => k1_relset_1(X0,X1) = k9_xtuple_0(X1)) [input redefinition_k1_relset_1]
25. ! [X0,X1] : ((v5_relat_1(X1,X0) & v1_relat_1(X1)) => k2_relset_1(X0,X1) = k10_xtuple_0(X1)) [input redefinition_k2_relset_1]
29. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
30. ! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1)) => r1_tarski(k9_xtuple_0(k2_funct_4(X2)),k2_zfmisc_1(X1,X0)))) [input t45_funct_4]
31. ! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1)) => k10_xtuple_0(k2_funct_4(X2)) = k10_xtuple_0(X2))) [input t47_funct_4]
32. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => ((r1_tarski(k10_xtuple_0(X2),X1) & r1_tarski(k9_xtuple_0(X2),X0)) => m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))))) [input t4_relset_1]
33. ~! [X0,X1,X2,X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2))) & v1_funct_1(X3)) => (m1_subset_1(k2_funct_4(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X0),X2))) & v1_funct_1(k2_funct_4(X3)))) [flattening 2]
34. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 29]
35. ! [X0,X1,X2] : (v1_relat_1(X2) => ((r1_tarski(k10_xtuple_0(X2),X1) & r1_tarski(k9_xtuple_0(X2),X0)) => m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))))) [flattening 32]
36. ! [X0,X1,X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1)) => k10_xtuple_0(k2_funct_4(X2)) = k10_xtuple_0(X2))) [flattening 31]
37. ! [X0,X1,X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1)) => r1_tarski(k9_xtuple_0(k2_funct_4(X2)),k2_zfmisc_1(X1,X0)))) [flattening 30]
40. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [flattening 5]
42. ? [X0,X1,X2,X3] : ((~m1_subset_1(k2_funct_4(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X0),X2))) | ~v1_funct_1(k2_funct_4(X3))) & (m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2))) & v1_funct_1(X3))) [ennf transformation 33]
43. ? [X0,X1,X2,X3] : ((~m1_subset_1(k2_funct_4(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X0),X2))) | ~v1_funct_1(k2_funct_4(X3))) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2))) & v1_funct_1(X3)) [flattening 42]
44. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
45. ! [X0] : ((v1_funct_1(k2_funct_4(X0)) & v1_relat_1(k2_funct_4(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 12]
46. ! [X0] : ((v1_funct_1(k2_funct_4(X0)) & v1_relat_1(k2_funct_4(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 45]
49. ! [X0,X1,X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | (~r1_tarski(k10_xtuple_0(X2),X1) | ~r1_tarski(k9_xtuple_0(X2),X0))) | ~v1_relat_1(X2)) [ennf transformation 35]
50. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~r1_tarski(k10_xtuple_0(X2),X1) | ~r1_tarski(k9_xtuple_0(X2),X0) | ~v1_relat_1(X2)) [flattening 49]
51. ! [X0,X1,X2] : ((k10_xtuple_0(k2_funct_4(X2)) = k10_xtuple_0(X2) | ~r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1))) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) [ennf transformation 36]
52. ! [X0,X1,X2] : (k10_xtuple_0(k2_funct_4(X2)) = k10_xtuple_0(X2) | ~r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1)) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) [flattening 51]
53. ! [X0,X1,X2] : ((r1_tarski(k9_xtuple_0(k2_funct_4(X2)),k2_zfmisc_1(X1,X0)) | ~r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1))) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) [ennf transformation 37]
54. ! [X0,X1,X2] : (r1_tarski(k9_xtuple_0(k2_funct_4(X2)),k2_zfmisc_1(X1,X0)) | ~r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1)) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) [flattening 53]
55. ! [X0,X1] : (k2_relset_1(X0,X1) = k10_xtuple_0(X1) | (~v5_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 25]
56. ! [X0,X1] : (k2_relset_1(X0,X1) = k10_xtuple_0(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 55]
57. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 24]
58. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 57]
59. ! [X0,X1] : (m1_subset_1(k1_relset_1(X0,X1),k1_zfmisc_1(X0)) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 10]
60. ! [X0,X1] : (m1_subset_1(k1_relset_1(X0,X1),k1_zfmisc_1(X0)) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 59]
61. ! [X0,X1] : (m1_subset_1(k2_relset_1(X0,X1),k1_zfmisc_1(X0)) | (~v5_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 13]
62. ! [X0,X1] : (m1_subset_1(k2_relset_1(X0,X1),k1_zfmisc_1(X0)) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 61]
68. ! [X0,X1,X2] : ((v5_relat_1(X2,X1) & v4_relat_1(X2,X0)) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 40]
73. ? [X0,X1,X2,X3] : ((~m1_subset_1(k2_funct_4(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X0),X2))) | ~v1_funct_1(k2_funct_4(X3))) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2))) & v1_funct_1(X3)) => ((~m1_subset_1(k2_funct_4(sK3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK1,sK0),sK2))) | ~v1_funct_1(k2_funct_4(sK3))) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK1),sK2))) & v1_funct_1(sK3)) [choice axiom]
74. (~m1_subset_1(k2_funct_4(sK3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK1,sK0),sK2))) | ~v1_funct_1(k2_funct_4(sK3))) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK1),sK2))) & v1_funct_1(sK3) [skolemisation 43,73]
75. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 34]
85. v1_funct_1(sK3) [cnf transformation 74] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
86. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK1),sK2))) [cnf transformation 74] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
87. ~m1_subset_1(k2_funct_4(sK3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK1,sK0),sK2))) | ~v1_funct_1(k2_funct_4(sK3)) [cnf transformation 74] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
88. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 44] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
89. v1_relat_1(k2_funct_4(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
90. v1_funct_1(k2_funct_4(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
92. ~m1_subset_1(X0,k1_zfmisc_1(X1)) | r1_tarski(X0,X1) [cnf transformation 75] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
94. m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | ~r1_tarski(k10_xtuple_0(X2),X1) | ~r1_tarski(k9_xtuple_0(X2),X0) | ~v1_relat_1(X2) [cnf transformation 50] {a:0,w:16,wCS:6400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
95. ~r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1)) | k10_xtuple_0(k2_funct_4(X2)) = k10_xtuple_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [cnf transformation 52] {a:0,w:16,wCS:6400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
96. r1_tarski(k9_xtuple_0(k2_funct_4(X2)),k2_zfmisc_1(X1,X0)) | ~r1_tarski(k9_xtuple_0(X2),k2_zfmisc_1(X0,X1)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [cnf transformation 54] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
99. v1_relat_1(k2_zfmisc_1(X0,X1)) [cnf transformation 18] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
100. k2_relset_1(X0,X1) = k10_xtuple_0(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 56] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
101. k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 58] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
102. m1_subset_1(k1_relset_1(X0,X1),k1_zfmisc_1(X0)) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 60] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
103. m1_subset_1(k2_relset_1(X0,X1),k1_zfmisc_1(X0)) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 62] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
107. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v4_relat_1(X2,X0) [cnf transformation 68] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
108. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v5_relat_1(X2,X1) [cnf transformation 68] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
124. 1 <=> v1_funct_1(k2_funct_4(sK3)) [avatar definition]
126. ~v1_funct_1(k2_funct_4(sK3)) <- (~1) [avatar component clause 124] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
128. 2 <=> m1_subset_1(k2_funct_4(sK3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK1,sK0),sK2))) [avatar definition]
130. ~m1_subset_1(k2_funct_4(sK3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK1,sK0),sK2))) <- (~2) [avatar component clause 128] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
131. ~1 | ~2 [avatar split clause 87,128,124]
133. 3 <=> m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK1),sK2))) [avatar definition]
135. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK1),sK2))) <- (3) [avatar component clause 133] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
136. 3 [avatar split clause 86,133]
138. 4 <=> v1_funct_1(sK3) [avatar definition]
140. v1_funct_1(sK3) <- (4) [avatar component clause 138] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
141. 4 [avatar split clause 85,138]
152. ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (~1) [resolution 90,126] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
153. ~v1_relat_1(sK3) <- (~1, 4) [subsumption resolution 152,140] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
155. 7 <=> v1_relat_1(sK3) [avatar definition]
156. v1_relat_1(sK3) <- (7) [avatar component clause 155] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
158. ~7 | 1 | ~4 [avatar split clause 153,138,124,155]
168. v1_relat_1(sK3) | ~v1_relat_1(k2_zfmisc_1(k2_zfmisc_1(sK0,sK1),sK2)) <- (3) [resolution 88,135] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
175. v1_relat_1(sK3) <- (3) [subsumption resolution 168,99] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
176. 7 | ~3 [avatar split clause 175,133,155]
186. v4_relat_1(sK3,k2_zfmisc_1(sK0,sK1)) <- (3) [resolution 107,135] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
190. 10 <=> v4_relat_1(sK3,k2_zfmisc_1(sK0,sK1)) [avatar definition]
192. v4_relat_1(sK3,k2_zfmisc_1(sK0,sK1)) <- (10) [avatar component clause 190] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
193. 10 | ~3 [avatar split clause 186,133,190]
194. v5_relat_1(sK3,sK2) <- (3) [resolution 108,135] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
198. 11 <=> v5_relat_1(sK3,sK2) [avatar definition]
200. v5_relat_1(sK3,sK2) <- (11) [avatar component clause 198] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
201. 11 | ~3 [avatar split clause 194,133,198]
207. r1_tarski(k1_relset_1(X5,X4),X5) | ~v1_relat_1(X4) | ~v4_relat_1(X4,X5) [resolution 102,92] {a:1,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
215. r1_tarski(k2_relset_1(X5,X4),X5) | ~v1_relat_1(X4) | ~v5_relat_1(X4,X5) [resolution 103,92] {a:1,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
280. ~r1_tarski(k10_xtuple_0(k2_funct_4(sK3)),sK2) | ~r1_tarski(k9_xtuple_0(k2_funct_4(sK3)),k2_zfmisc_1(sK1,sK0)) | ~v1_relat_1(k2_funct_4(sK3)) <- (~2) [resolution 94,130] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
285. 14 <=> v1_relat_1(k2_funct_4(sK3)) [avatar definition]
287. ~v1_relat_1(k2_funct_4(sK3)) <- (~14) [avatar component clause 285] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
289. 15 <=> r1_tarski(k9_xtuple_0(k2_funct_4(sK3)),k2_zfmisc_1(sK1,sK0)) [avatar definition]
291. ~r1_tarski(k9_xtuple_0(k2_funct_4(sK3)),k2_zfmisc_1(sK1,sK0)) <- (~15) [avatar component clause 289] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
293. 16 <=> r1_tarski(k10_xtuple_0(k2_funct_4(sK3)),sK2) [avatar definition]
295. ~r1_tarski(k10_xtuple_0(k2_funct_4(sK3)),sK2) <- (~16) [avatar component clause 293] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
296. ~14 | ~15 | ~16 | 2 [avatar split clause 280,128,293,289,285]
298. ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (~14) [resolution 287,89] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
300. ~v1_relat_1(sK3) <- (4, ~14) [subsumption resolution 298,140] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
301. $false <- (4, 7, ~14) [subsumption resolution 300,156] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
302. ~4 | ~7 | 14 [avatar contradiction clause 301]
306. ~r1_tarski(k9_xtuple_0(sK3),k2_zfmisc_1(sK0,sK1)) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (~15) [resolution 291,96] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
307. ~r1_tarski(k9_xtuple_0(sK3),k2_zfmisc_1(sK0,sK1)) | ~v1_relat_1(sK3) <- (4, ~15) [subsumption resolution 306,140] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
308. ~r1_tarski(k9_xtuple_0(sK3),k2_zfmisc_1(sK0,sK1)) <- (4, 7, ~15) [subsumption resolution 307,156] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
310. 17 <=> r1_tarski(k9_xtuple_0(sK3),k2_zfmisc_1(sK0,sK1)) [avatar definition]
311. r1_tarski(k9_xtuple_0(sK3),k2_zfmisc_1(sK0,sK1)) <- (17) [avatar component clause 310] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
312. ~r1_tarski(k9_xtuple_0(sK3),k2_zfmisc_1(sK0,sK1)) <- (~17) [avatar component clause 310] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
313. ~17 | ~4 | ~7 | 15 [avatar split clause 308,289,155,138,310]
345. r1_tarski(k9_xtuple_0(X1),X0) | ~v1_relat_1(X1) | ~v4_relat_1(X1,X0) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [superposition 207,101] {a:2,w:14,wCS:5600,sine:2,thAx:0,allAx:3,thDist:-3}
346. r1_tarski(k9_xtuple_0(X1),X0) | ~v1_relat_1(X1) | ~v4_relat_1(X1,X0) [duplicate literal removal 345] {a:2,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:3,thDist:-3}
363. ~v1_relat_1(sK3) | ~v4_relat_1(sK3,k2_zfmisc_1(sK0,sK1)) <- (~17) [resolution 346,312] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
372. ~v4_relat_1(sK3,k2_zfmisc_1(sK0,sK1)) <- (7, ~17) [subsumption resolution 363,156] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
373. $false <- (7, 10, ~17) [subsumption resolution 372,192] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
374. ~7 | ~10 | 17 [avatar contradiction clause 373]
420. r1_tarski(k10_xtuple_0(X1),X0) | ~v1_relat_1(X1) | ~v5_relat_1(X1,X0) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [superposition 215,100] {a:2,w:14,wCS:5600,sine:2,thAx:0,allAx:3,thDist:-3}
421. r1_tarski(k10_xtuple_0(X1),X0) | ~v1_relat_1(X1) | ~v5_relat_1(X1,X0) [duplicate literal removal 420] {a:2,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:3,thDist:-3}
461. k10_xtuple_0(k2_funct_4(sK3)) = k10_xtuple_0(sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (17) [resolution 311,95] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
468. k10_xtuple_0(k2_funct_4(sK3)) = k10_xtuple_0(sK3) | ~v1_relat_1(sK3) <- (4, 17) [subsumption resolution 461,140] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
469. k10_xtuple_0(k2_funct_4(sK3)) = k10_xtuple_0(sK3) <- (4, 7, 17) [subsumption resolution 468,156] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
471. 25 <=> k10_xtuple_0(k2_funct_4(sK3)) = k10_xtuple_0(sK3) [avatar definition]
473. k10_xtuple_0(k2_funct_4(sK3)) = k10_xtuple_0(sK3) <- (25) [avatar component clause 471] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
474. 25 | ~4 | ~7 | ~17 [avatar split clause 469,310,155,138,471]
530. ~r1_tarski(k10_xtuple_0(sK3),sK2) <- (~16, 25) [backward demodulation 295,473] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
533. 31 <=> r1_tarski(k10_xtuple_0(sK3),sK2) [avatar definition]
535. ~r1_tarski(k10_xtuple_0(sK3),sK2) <- (~31) [avatar component clause 533] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
536. ~31 | 16 | ~25 [avatar split clause 530,471,293,533]
863. ~v1_relat_1(sK3) | ~v5_relat_1(sK3,sK2) <- (~31) [resolution 421,535] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
872. ~v5_relat_1(sK3,sK2) <- (7, ~31) [subsumption resolution 863,156] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
873. $false <- (7, 11, ~31) [subsumption resolution 872,200] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
874. ~7 | ~11 | 31 [avatar contradiction clause 873]
876. $false [avatar sat refutation 131,136,141,158,176,193,201,296,302,313,374,474,536,874] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:490,thDist:-490}
% SZS output end Proof for funct_4__t48_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 31
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 38
% Generated clauses: 590
% Active clauses: 170
% Passive clauses: 304
% Final active clauses: 159
% Final passive clauses: 129
% Main loop iterations started: 175
% 
% >>> Simplifying Inferences
% Duplicate literals: 74
% Fw subsumption resolutions: 61
% Fw demodulations: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 30
% Forward subsumptions: 69
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 332
% Unit resulting resolution: 10
% Forward superposition: 38
% Backward superposition: 7
% 
% >>> AVATAR
% Split clauses: 14
% Split components: 29
% Sat splitting refutations: 19
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 63
% SAT solver unit clauses: 9
% SAT solver binary clauses: 22
% 
% Memory used [KB]: 5756
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% backward subsumption index maintenance: 0.001 s
% backward superposition index maintenance: 0.001 s
% resolution: 0.001 s
% unit resulting resolution: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.009 s ( own 0.003 s ) 

