i: [1,24,0,2,1,1,0,0] t21_relat_1
a: [3,42,0,2,1,24]
s: 42
i: [1,25,0,6,1,1,0,0] t21_relat_1
a: [3,37,0,6,1,25]
s: 37
i: [1,26,0,5,2,0,0,1] dt_k2_relat_1
s: 26
i: [1,31,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 31
i: [1,27,0,7,2,0,0,1] involutiveness_k2_relat_1
s: 27
i: [1,32,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 32
i: [1,29,0,8,2,0,0,2] t20_relat_1
s: 29
i: [1,30,0,8,2,0,0,2] t20_relat_1
s: 30
i: [1,33,0,8,3,0,0,2] fc3_relat_1
s: 33
i: [1,28,0,10,2,0,0,1] d6_relat_1
s: 28
d: [2,45,1,11,3,0,44,29,27]
d: [2,46,1,8,2,0,21,45,26]
s: 46
d: [2,54,1,11,4,0,44,33,28]
s: 54
d: [2,52,1,15,3,0,44,28,29]
d: [2,55,1,12,2,0,21,52,26]
s: 55
d: [2,61,2,14,3,0,44,55,27]
d: [2,66,2,11,2,0,21,61,26]
d: [2,67,2,11,2,0,22,66,32]
s: 67
d: [2,53,1,15,3,0,44,28,30]
d: [2,56,1,12,2,0,21,53,26]
d: [2,57,1,12,2,0,22,56,32]
s: 57
d: [2,62,2,13,3,0,44,55,30]
d: [2,64,2,11,2,0,18,62]
d: [2,68,2,11,2,0,22,64,32]
s: 68
d: [2,73,3,10,3,0,44,55,67]
d: [2,75,3,8,2,0,18,73]
s: 75
d: [2,100,4,0,0,2,52,42,37,75]
e: 100
d: [2,102,4,7,2,1,44,37,75]
d: [2,103,4,2,1,1,19,102]
d: [2,105,4,0,0,2,21,103,42]
e: 105
f: 105,24,25
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t21_relat_1
% SZS output start Proof for relat_1__t21_relat_1
1. ! [X0] : (v1_relat_1(X0) => k1_relat_1(X0) = k1_relat_1(k2_relat_1(X0))) [input t21_relat_1]
2. ~! [X0] : (v1_relat_1(X0) => k1_relat_1(X0) = k1_relat_1(k2_relat_1(X0))) [negated conjecture 1]
3. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
4. ! [X0] : (v1_relat_1(X0) => k1_relat_1(X0) = k2_xboole_0(k9_xtuple_0(X0),k10_xtuple_0(X0))) [input d6_relat_1]
7. ! [X0] : (v1_relat_1(X0) => v1_relat_1(k2_relat_1(X0))) [input dt_k2_relat_1]
12. ! [X0] : (v1_relat_1(X0) => k2_relat_1(k2_relat_1(X0)) = X0) [input involutiveness_k2_relat_1]
13. ! [X0] : (v1_relat_1(X0) => (k9_xtuple_0(X0) = k10_xtuple_0(k2_relat_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_relat_1(X0)))) [input t20_relat_1]
15. ? [X0] : (k1_relat_1(X0) != k1_relat_1(k2_relat_1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
16. ! [X0] : (v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 7]
17. ! [X0] : (k2_relat_1(k2_relat_1(X0)) = X0 | ~v1_relat_1(X0)) [ennf transformation 12]
18. ! [X0] : (k1_relat_1(X0) = k2_xboole_0(k9_xtuple_0(X0),k10_xtuple_0(X0)) | ~v1_relat_1(X0)) [ennf transformation 4]
19. ! [X0] : ((k9_xtuple_0(X0) = k10_xtuple_0(k2_relat_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 13]
22. ? [X0] : (k1_relat_1(X0) != k1_relat_1(k2_relat_1(X0)) & v1_relat_1(X0)) => (k1_relat_1(sK0) != k1_relat_1(k2_relat_1(sK0)) & v1_relat_1(sK0)) [choice axiom]
23. k1_relat_1(sK0) != k1_relat_1(k2_relat_1(sK0)) & v1_relat_1(sK0) [skolemisation 15,22]
24. v1_relat_1(sK0) [cnf transformation 23] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
25. k1_relat_1(sK0) != k1_relat_1(k2_relat_1(sK0)) [cnf transformation 23] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
26. v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 16] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
27. k2_relat_1(k2_relat_1(X0)) = X0 | ~v1_relat_1(X0) [cnf transformation 17] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
28. k1_relat_1(X0) = k2_xboole_0(k9_xtuple_0(X0),k10_xtuple_0(X0)) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
29. k10_xtuple_0(X0) = k9_xtuple_0(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 19] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
32. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
35. 1 <=> k1_relat_1(sK0) = k1_relat_1(k2_relat_1(sK0)) [avatar definition]
37. k1_relat_1(sK0) != k1_relat_1(k2_relat_1(sK0)) <- (~1) [avatar component clause 35] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
38. ~1 [avatar split clause 25,35]
40. 2 <=> v1_relat_1(sK0) [avatar definition]
42. v1_relat_1(sK0) <- (2) [avatar component clause 40] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
43. 2 [avatar split clause 24,40]
52. k1_relat_1(k2_relat_1(X0)) = k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(k2_relat_1(X0))) | ~v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [superposition 28,29] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
55. k1_relat_1(k2_relat_1(X0)) = k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(k2_relat_1(X0))) | ~v1_relat_1(X0) [subsumption resolution 52,26] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
61. k1_relat_1(X0) = k2_xboole_0(k10_xtuple_0(k2_relat_1(X0)),k10_xtuple_0(X0)) | ~v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [superposition 55,27] {a:2,w:14,wCS:5600,sine:1,thAx:0,allAx:3,thDist:-3}
66. k1_relat_1(X0) = k2_xboole_0(k10_xtuple_0(k2_relat_1(X0)),k10_xtuple_0(X0)) | ~v1_relat_1(X0) [subsumption resolution 61,26] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:3,thDist:-3}
67. k1_relat_1(X0) = k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(k2_relat_1(X0))) | ~v1_relat_1(X0) [forward demodulation 66,32] {a:2,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
73. k1_relat_1(X0) = k1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_1(X0) [superposition 55,67] {a:3,w:10,wCS:4000,sine:1,thAx:0,allAx:5,thDist:-5}
75. k1_relat_1(X0) = k1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [duplicate literal removal 73] {a:3,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
102. k1_relat_1(sK0) != k1_relat_1(sK0) | ~v1_relat_1(sK0) <- (~1) [superposition 37,75] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
103. ~v1_relat_1(sK0) <- (~1) [trivial inequality removal 102] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
105. $false <- (~1, 2) [subsumption resolution 103,42] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
106. 1 | ~2 [avatar contradiction clause 105]
107. $false [avatar sat refutation 38,43,106] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:22,thDist:-22}
% SZS output end Proof for relat_1__t21_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 1
% 
% >>> Saturation
% Initial clauses: 10
% Generated clauses: 73
% Active clauses: 16
% Passive clauses: 32
% Final active clauses: 16
% Final passive clauses: 15
% Main loop iterations started: 17
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Trivial inequalities: 1
% Fw subsumption resolutions: 13
% Fw demodulations: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Unit resulting resolution: 1
% Forward superposition: 22
% Backward superposition: 9
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 6
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

