i: [1,53,0,3,1,1,0,0] l18_zfmisc_1
a: [3,89,0,3,1,53]
s: 89
i: [1,52,0,4,1,1,0,0] l18_zfmisc_1
a: [3,94,0,4,1,52]
s: 94
i: [1,54,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,99,0,2,1,54]
s: 99
i: [1,55,0,2,1,0,0,2] fc1_xboole_0
a: [3,104,0,2,1,55]
s: 104
i: [1,78,0,2,1,0,0,2] rc2_xboole_0
a: [3,114,0,2,1,78]
s: 114
i: [1,79,0,2,1,0,0,2] rc1_xboole_0
a: [3,119,0,2,1,79]
s: 119
i: [1,57,0,3,1,0,0,1] fc2_xboole_0
s: 57
i: [1,56,0,3,1,0,0,255] d2_xboole_0
a: [3,109,0,3,1,56]
s: 109
r: 99
i: [1,81,0,4,1,0,0,1] d1_tarski
s: 81
i: [1,58,0,5,1,0,0,2] t2_boole
s: 58
i: [1,59,0,5,2,0,0,2] t6_boole
s: 59
i: [1,67,0,5,2,0,0,1] t7_boole
s: 67
i: [1,74,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 74
d: [2,125,1,2,1,1,40,67,89]
a: [3,130,1,2,1,125]
s: 130
i: [1,76,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 76
i: [1,77,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 77
d: [2,132,1,3,1,1,40,76,89]
a: [3,137,1,3,1,132]
s: 137
i: [1,66,0,7,3,0,0,2] t8_boole
s: 66
d: [2,139,1,4,1,1,40,77,94]
a: [3,143,1,4,1,139]
s: 143
i: [1,75,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 75
i: [1,82,0,7,2,0,0,1] d1_tarski
s: 82
d: [2,133,1,4,1,0,40,76,81]
s: 133
i: [1,60,0,8,2,0,0,1] d7_xboole_0
s: 60
d: [2,146,1,5,1,0,44,75,58]
s: 146
i: [1,61,0,8,2,0,0,1] d7_xboole_0
s: 61
i: [1,84,0,8,2,0,0,1] d4_xboole_0
s: 84
d: [2,162,1,6,2,0,44,61,58]
d: [2,169,1,3,1,0,19,162]
s: 169
i: [1,85,0,8,2,0,0,1] d4_xboole_0
s: 85
d: [2,167,2,6,2,0,44,61,146]
d: [2,168,2,3,1,0,19,167]
s: 168
i: [1,83,0,11,3,0,0,1] d4_xboole_0
s: 83
i: [1,64,0,14,3,0,0,1] d1_tarski
s: 64
d: [2,176,2,3,1,0,40,169,77]
s: 176
i: [1,65,0,14,3,0,0,1] d1_tarski
s: 65
d: [2,147,1,5,1,0,44,75,58]
s: 147
i: [1,71,0,17,3,0,0,1] d4_xboole_0
s: 71
i: [1,72,0,17,3,0,0,1] d4_xboole_0
s: 72
d: [2,148,1,5,1,0,44,58,75]
s: 148
i: [1,73,0,23,4,0,0,1] d4_xboole_0
s: 73
d: [2,149,1,5,1,0,44,58,75]
s: 149
d: [2,151,1,6,2,0,44,60,74]
s: 151
d: [2,154,1,6,2,0,44,74,60]
s: 154
d: [2,177,2,5,2,0,44,169,59]
s: 177
d: [2,163,1,6,2,0,44,61,74]
s: 163
d: [2,185,3,5,2,0,44,168,59]
s: 185
d: [2,178,1,6,2,0,44,85,58]
s: 178
d: [2,122,1,7,3,0,44,59,59]
s: 122
d: [2,222,3,5,2,0,40,177,77]
s: 222
d: [2,123,1,7,2,0,44,58,59]
s: 123
d: [2,175,2,6,2,0,44,84,146]
s: 175
d: [2,152,1,8,2,0,44,60,75]
s: 152
d: [2,153,1,8,2,0,44,60,75]
s: 153
d: [2,157,2,7,2,0,44,146,59]
s: 157
d: [2,155,1,8,2,0,44,75,60]
s: 155
d: [2,233,2,7,2,0,44,123,75]
s: 233
d: [2,156,1,8,2,0,44,75,60]
s: 156
d: [2,165,1,8,2,0,44,61,75]
s: 165
d: [2,234,2,7,2,0,44,123,75]
s: 234
d: [2,166,1,8,2,0,44,61,75]
s: 166
d: [2,238,2,7,2,0,44,75,123]
s: 238
d: [2,173,1,8,2,0,44,84,75]
s: 173
d: [2,174,1,8,2,0,44,84,75]
s: 174
d: [2,239,2,7,2,0,44,75,123]
s: 239
d: [2,172,1,9,3,0,44,84,60]
s: 172
d: [2,224,2,8,3,0,44,163,59]
s: 224
d: [2,180,1,9,3,0,44,85,60]
s: 180
d: [2,189,1,10,3,0,40,83,67]
s: 189
d: [2,289,2,11,4,0,44,189,152]
d: [2,298,2,9,3,1,21,289,104]
s: 298
d: [2,300,3,0,0,3,52,143,81,89,298]
e: 300
d: [2,301,3,0,0,3,52,89,94,81,298]
e: 301
f: 301,55,52,53
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__l18_zfmisc_1
% SZS output start Proof for zfmisc_1__l18_zfmisc_1
1. ! [X0] : ! [X1] : ~(r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [input l18_zfmisc_1]
2. ~! [X0] : ! [X1] : ~(r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [negated conjecture 1]
4. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
5. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
7. ! [X0] : ! [X1] : ! [X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [input d4_xboole_0]
8. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) <=> k3_xboole_0(X0,X1) = k1_xboole_0) [input d7_xboole_0]
13. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
21. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
23. ~! [X0,X1] : ~(r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [flattening 2]
24. ! [X0,X1] : (r1_xboole_0(X0,X1) <=> k3_xboole_0(X0,X1) = k1_xboole_0) [flattening 8]
25. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 5]
27. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 21]
28. ! [X0,X1,X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [flattening 7]
30. ? [X0,X1] : (r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [ennf transformation 23]
33. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 27]
36. ? [X0,X1] : (r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) => (r2_hidden(sK0,sK1) & r1_xboole_0(k1_tarski(sK0),sK1)) [choice axiom]
37. r2_hidden(sK0,sK1) & r1_xboole_0(k1_tarski(sK0),sK1) [skolemisation 30,36]
38. ! [X0,X1] : ((r1_xboole_0(X0,X1) | k3_xboole_0(X0,X1) != k1_xboole_0) & (k3_xboole_0(X0,X1) = k1_xboole_0 | ~r1_xboole_0(X0,X1))) [nnf transformation 24]
39. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 25]
40. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 39]
41. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) [choice axiom]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 40,41]
43. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) | ~r2_hidden(X3,X0))) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [nnf transformation 28]
44. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | ~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [flattening 43]
45. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [rectify 44]
46. ! [X2,X1,X0] : (? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2))) => ((~r2_hidden(sK3(X0,X1,X2),X1) | ~r2_hidden(sK3(X0,X1,X2),X0) | ~r2_hidden(sK3(X0,X1,X2),X2)) & ((r2_hidden(sK3(X0,X1,X2),X1) & r2_hidden(sK3(X0,X1,X2),X0)) | r2_hidden(sK3(X0,X1,X2),X2)))) [choice axiom]
47. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ((~r2_hidden(sK3(X0,X1,X2),X1) | ~r2_hidden(sK3(X0,X1,X2),X0) | ~r2_hidden(sK3(X0,X1,X2),X2)) & ((r2_hidden(sK3(X0,X1,X2),X1) & r2_hidden(sK3(X0,X1,X2),X0)) | r2_hidden(sK3(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [skolemisation 45,46]
52. r1_xboole_0(k1_tarski(sK0),sK1) [cnf transformation 37] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. r2_hidden(sK0,sK1) [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. v1_xboole_0(k1_xboole_0) [cnf transformation 13] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
60. k3_xboole_0(X0,X1) = k1_xboole_0 | ~r1_xboole_0(X0,X1) [cnf transformation 38] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 42] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
67. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 33] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) | k3_xboole_0(X0,X1) != X2 [cnf transformation 47] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
75. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
80. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 63] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:2,thDist:-2}
81. r2_hidden(X3,k1_tarski(X3)) [equality resolution 80] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
83. r2_hidden(X4,k3_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) [equality resolution 70] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
87. 1 <=> r2_hidden(sK0,sK1) [avatar definition]
89. r2_hidden(sK0,sK1) <- (1) [avatar component clause 87] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
90. 1 [avatar split clause 53,87]
92. 2 <=> r1_xboole_0(k1_tarski(sK0),sK1) [avatar definition]
94. r1_xboole_0(k1_tarski(sK0),sK1) <- (2) [avatar component clause 92] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
95. 2 [avatar split clause 52,92]
102. 4 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
104. v1_xboole_0(k1_xboole_0) <- (4) [avatar component clause 102] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
105. 4 [avatar split clause 55,102]
152. k1_xboole_0 = k3_xboole_0(X5,X4) | ~r1_xboole_0(X4,X5) [superposition 60,75] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
189. ~v1_xboole_0(k3_xboole_0(X11,X10)) | ~r2_hidden(X9,X11) | ~r2_hidden(X9,X10) [resolution 83,67] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
289. ~v1_xboole_0(k1_xboole_0) | ~r2_hidden(X9,X7) | ~r2_hidden(X9,X8) | ~r1_xboole_0(X8,X7) [superposition 189,152] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:5,thDist:-5}
298. ~r2_hidden(X9,X8) | ~r2_hidden(X9,X7) | ~r1_xboole_0(X8,X7) <- (4) [subsumption resolution 289,104] {a:2,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
301. $false <- (1, 2, 4) [unit resulting resolution 89,94,81,298] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
309. ~1 | ~2 | ~4 [avatar contradiction clause 301]
310. $false [avatar sat refutation 90,95,105,309] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:56,thDist:-56}
% SZS output end Proof for zfmisc_1__l18_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 28
% Generated clauses: 210
% Active clauses: 47
% Passive clauses: 93
% Final active clauses: 46
% Final passive clauses: 26
% Main loop iterations started: 67
% 
% >>> Simplifying Inferences
% Duplicate literals: 6
% Trivial inequalities: 6
% Fw subsumption resolutions: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Equational tautologies: 3
% Forward subsumptions: 84
% 
% >>> Generating Inferences
% Binary resolution: 47
% Unit resulting resolution: 2
% Factoring: 2
% Forward superposition: 68
% Backward superposition: 35
% Self superposition: 1
% Equality resolution: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 7
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

