i: [1,27,0,3,1,1,0,0] t110_xboole_1
a: [3,44,0,3,1,27]
s: 44
i: [1,26,0,3,1,1,0,0] t110_xboole_1
a: [3,49,0,3,1,26]
s: 49
i: [1,28,0,5,1,1,0,0] t110_xboole_1
a: [3,39,0,5,1,28]
s: 39
i: [1,32,0,3,1,0,0,1] reflexivity_r1_tarski
s: 32
i: [1,33,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 33
i: [1,34,0,7,1,0,0,1] commutativity_k5_xboole_0
s: 34
i: [1,35,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 35
i: [1,30,0,8,2,0,0,255] t109_xboole_1
s: 30
i: [1,29,0,11,1,0,0,255] d6_xboole_0
s: 29
i: [1,31,0,11,3,0,0,1] t8_xboole_1
s: 31
d: [2,51,1,7,1,0,44,29,33]
s: 51
d: [2,54,1,7,1,0,44,33,29]
s: 54
d: [2,52,1,11,1,0,44,29,35]
s: 52
d: [2,64,2,8,2,0,44,30,51]
s: 64
d: [2,53,1,11,1,0,44,29,35]
s: 53
d: [2,55,1,11,1,0,44,35,29]
s: 55
d: [2,56,1,11,1,0,44,35,29]
s: 56
d: [2,60,1,15,3,0,44,31,29]
s: 60
d: [2,77,2,13,3,0,40,60,30]
s: 77
d: [2,76,2,14,2,0,40,60,32]
s: 76
d: [2,73,2,15,3,0,44,31,52]
s: 73
d: [2,81,3,11,3,0,40,77,30]
s: 81
d: [2,88,4,0,0,3,52,49,44,39,81]
e: 88
d: [2,89,4,6,2,1,40,81,39]
d: [2,93,4,3,1,2,21,89,44]
d: [2,94,4,0,0,3,21,93,49]
e: 94
f: 94,26,27,28
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t110_xboole_1
% SZS output start Proof for xboole_1__t110_xboole_1
1. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k5_xboole_0(X0,X2),X1)) [input t110_xboole_1]
2. ~! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k5_xboole_0(X0,X2),X1)) [negated conjecture 1]
5. ! [X0] : ! [X1] : k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [input d6_xboole_0]
11. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,X1) => r1_tarski(k4_xboole_0(X0,X2),X1)) [input t109_xboole_1]
12. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [input t8_xboole_1]
13. ~! [X0,X1,X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k5_xboole_0(X0,X2),X1)) [flattening 2]
14. ! [X0,X1] : k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [flattening 5]
15. ! [X0,X1,X2] : (r1_tarski(X0,X1) => r1_tarski(k4_xboole_0(X0,X2),X1)) [flattening 11]
16. ! [X0,X1,X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [flattening 12]
19. ? [X0,X1,X2] : (~r1_tarski(k5_xboole_0(X0,X2),X1) & (r1_tarski(X2,X1) & r1_tarski(X0,X1))) [ennf transformation 13]
20. ? [X0,X1,X2] : (~r1_tarski(k5_xboole_0(X0,X2),X1) & r1_tarski(X2,X1) & r1_tarski(X0,X1)) [flattening 19]
21. ! [X0,X1,X2] : (r1_tarski(k4_xboole_0(X0,X2),X1) | ~r1_tarski(X0,X1)) [ennf transformation 15]
22. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | (~r1_tarski(X2,X1) | ~r1_tarski(X0,X1))) [ennf transformation 16]
23. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1)) [flattening 22]
24. ? [X0,X1,X2] : (~r1_tarski(k5_xboole_0(X0,X2),X1) & r1_tarski(X2,X1) & r1_tarski(X0,X1)) => (~r1_tarski(k5_xboole_0(sK0,sK2),sK1) & r1_tarski(sK2,sK1) & r1_tarski(sK0,sK1)) [choice axiom]
25. ~r1_tarski(k5_xboole_0(sK0,sK2),sK1) & r1_tarski(sK2,sK1) & r1_tarski(sK0,sK1) [skolemisation 20,24]
26. r1_tarski(sK0,sK1) [cnf transformation 25] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. r1_tarski(sK2,sK1) [cnf transformation 25] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. ~r1_tarski(k5_xboole_0(sK0,sK2),sK1) [cnf transformation 25] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [cnf transformation 14] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
30. r1_tarski(k4_xboole_0(X0,X2),X1) | ~r1_tarski(X0,X1) [cnf transformation 21] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
31. r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1) [cnf transformation 23] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. 1 <=> r1_tarski(k5_xboole_0(sK0,sK2),sK1) [avatar definition]
39. ~r1_tarski(k5_xboole_0(sK0,sK2),sK1) <- (~1) [avatar component clause 37] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
40. ~1 [avatar split clause 28,37]
42. 2 <=> r1_tarski(sK2,sK1) [avatar definition]
44. r1_tarski(sK2,sK1) <- (2) [avatar component clause 42] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
45. 2 [avatar split clause 27,42]
47. 3 <=> r1_tarski(sK0,sK1) [avatar definition]
49. r1_tarski(sK0,sK1) <- (3) [avatar component clause 47] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
50. 3 [avatar split clause 26,47]
60. ~r1_tarski(k4_xboole_0(X9,X8),X10) | r1_tarski(k5_xboole_0(X8,X9),X10) | ~r1_tarski(k4_xboole_0(X8,X9),X10) [superposition 31,29] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
77. ~r1_tarski(k4_xboole_0(X2,X3),X4) | r1_tarski(k5_xboole_0(X2,X3),X4) | ~r1_tarski(X3,X4) [resolution 60,30] {a:2,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
81. r1_tarski(k5_xboole_0(X2,X3),X4) | ~r1_tarski(X3,X4) | ~r1_tarski(X2,X4) [resolution 77,30] {a:3,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
89. ~r1_tarski(sK2,sK1) | ~r1_tarski(sK0,sK1) <- (~1) [resolution 81,39] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
93. ~r1_tarski(sK0,sK1) <- (~1, 2) [subsumption resolution 89,44] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
94. $false <- (~1, 2, 3) [subsumption resolution 93,49] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
95. 1 | ~2 | ~3 [avatar contradiction clause 94]
96. $false [avatar sat refutation 40,45,50,95] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:25,thDist:-25}
% SZS output end Proof for xboole_1__t110_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 10
% Generated clauses: 56
% Active clauses: 17
% Passive clauses: 24
% Final active clauses: 17
% Final passive clauses: 2
% Main loop iterations started: 22
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 2
% Fw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 22
% 
% >>> Generating Inferences
% Binary resolution: 7
% Unit resulting resolution: 1
% Forward superposition: 19
% Backward superposition: 11
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 7
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

