i: [1,47,0,2,1,1,0,0] t87_orders_1
a: [3,64,0,2,1,47]
s: 64
i: [1,44,0,2,1,1,0,0] t87_orders_1
a: [3,79,0,2,1,44]
s: 79
i: [1,46,0,3,1,1,0,0] t87_orders_1
a: [3,69,0,3,1,46]
s: 69
i: [1,45,0,3,1,1,0,0] t87_orders_1
a: [3,74,0,3,1,45]
s: 74
i: [1,48,0,4,2,0,0,2] cc5_finset_1
s: 48
i: [1,49,0,4,2,0,0,2] cc4_finset_1
s: 49
i: [1,57,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 57
i: [1,51,0,7,3,0,0,255] t86_orders_1
s: 51
i: [1,52,0,7,3,0,0,1] fc25_relat_1
s: 52
i: [1,53,0,7,3,0,0,1] fc24_relat_1
s: 53
i: [1,54,0,7,3,0,0,255] fc28_finset_1
s: 54
i: [1,55,0,7,3,0,0,1] fc19_finset_1
s: 55
i: [1,56,0,7,3,0,0,1] fc22_finset_1
s: 56
i: [1,58,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 58
i: [1,59,0,8,3,0,0,255] fc3_relat_1
s: 59
i: [1,60,0,8,3,0,0,255] fc9_finset_1
s: 60
d: [2,81,1,7,3,0,40,52,48]
s: 81
i: [1,50,0,10,2,0,0,255] d6_relat_1
s: 50
d: [2,82,1,7,3,0,40,53,48]
s: 82
d: [2,93,1,11,4,0,44,60,50]
s: 93
d: [2,94,1,11,4,0,44,59,50]
s: 94
d: [2,95,2,8,3,1,40,93,69]
d: [2,100,2,5,2,2,21,95,74]
d: [2,101,2,3,1,3,21,100,79]
a: [3,105,2,3,1,101]
s: 105
d: [2,108,3,0,0,3,52,79,64,105,51]
e: 108
d: [2,109,3,4,2,1,40,105,51]
d: [2,111,3,2,1,2,21,109,64]
d: [2,112,3,0,0,3,21,111,79]
e: 112
f: 112,101,44,45,46,47
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t87_orders_1
% SZS output start Proof for orders_1__t87_orders_1
1. ! [X0] : (v1_relat_1(X0) => ((v1_finset_1(k10_xtuple_0(X0)) & v1_finset_1(k9_xtuple_0(X0))) => v1_finset_1(X0))) [input t87_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => ((v1_finset_1(k10_xtuple_0(X0)) & v1_finset_1(k9_xtuple_0(X0))) => v1_finset_1(X0))) [negated conjecture 1]
6. ! [X0] : (v1_relat_1(X0) => k1_relat_1(X0) = k2_xboole_0(k9_xtuple_0(X0),k10_xtuple_0(X0))) [input d6_relat_1]
17. ! [X0,X1] : ((v1_finset_1(X1) & v1_finset_1(X0)) => v1_finset_1(k2_xboole_0(X0,X1))) [input fc9_finset_1]
19. ! [X0] : (v1_relat_1(X0) => (v1_finset_1(k1_relat_1(X0)) => v1_finset_1(X0))) [input t86_orders_1]
21. ? [X0] : ((~v1_finset_1(X0) & (v1_finset_1(k10_xtuple_0(X0)) & v1_finset_1(k9_xtuple_0(X0)))) & v1_relat_1(X0)) [ennf transformation 2]
22. ? [X0] : (~v1_finset_1(X0) & v1_finset_1(k10_xtuple_0(X0)) & v1_finset_1(k9_xtuple_0(X0)) & v1_relat_1(X0)) [flattening 21]
25. ! [X0] : (k1_relat_1(X0) = k2_xboole_0(k9_xtuple_0(X0),k10_xtuple_0(X0)) | ~v1_relat_1(X0)) [ennf transformation 6]
26. ! [X0] : ((v1_finset_1(X0) | ~v1_finset_1(k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 19]
27. ! [X0] : (v1_finset_1(X0) | ~v1_finset_1(k1_relat_1(X0)) | ~v1_relat_1(X0)) [flattening 26]
40. ! [X0,X1] : (v1_finset_1(k2_xboole_0(X0,X1)) | (~v1_finset_1(X1) | ~v1_finset_1(X0))) [ennf transformation 17]
41. ! [X0,X1] : (v1_finset_1(k2_xboole_0(X0,X1)) | ~v1_finset_1(X1) | ~v1_finset_1(X0)) [flattening 40]
42. ? [X0] : (~v1_finset_1(X0) & v1_finset_1(k10_xtuple_0(X0)) & v1_finset_1(k9_xtuple_0(X0)) & v1_relat_1(X0)) => (~v1_finset_1(sK0) & v1_finset_1(k10_xtuple_0(sK0)) & v1_finset_1(k9_xtuple_0(sK0)) & v1_relat_1(sK0)) [choice axiom]
43. ~v1_finset_1(sK0) & v1_finset_1(k10_xtuple_0(sK0)) & v1_finset_1(k9_xtuple_0(sK0)) & v1_relat_1(sK0) [skolemisation 22,42]
44. v1_relat_1(sK0) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. v1_finset_1(k9_xtuple_0(sK0)) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. v1_finset_1(k10_xtuple_0(sK0)) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. ~v1_finset_1(sK0) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. k1_relat_1(X0) = k2_xboole_0(k9_xtuple_0(X0),k10_xtuple_0(X0)) | ~v1_relat_1(X0) [cnf transformation 25] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
51. ~v1_finset_1(k1_relat_1(X0)) | v1_finset_1(X0) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
60. v1_finset_1(k2_xboole_0(X0,X1)) | ~v1_finset_1(X1) | ~v1_finset_1(X0) [cnf transformation 41] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
62. 1 <=> v1_finset_1(sK0) [avatar definition]
64. ~v1_finset_1(sK0) <- (~1) [avatar component clause 62] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. ~1 [avatar split clause 47,62]
67. 2 <=> v1_finset_1(k10_xtuple_0(sK0)) [avatar definition]
69. v1_finset_1(k10_xtuple_0(sK0)) <- (2) [avatar component clause 67] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
70. 2 [avatar split clause 46,67]
72. 3 <=> v1_finset_1(k9_xtuple_0(sK0)) [avatar definition]
74. v1_finset_1(k9_xtuple_0(sK0)) <- (3) [avatar component clause 72] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. 3 [avatar split clause 45,72]
77. 4 <=> v1_relat_1(sK0) [avatar definition]
79. v1_relat_1(sK0) <- (4) [avatar component clause 77] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. 4 [avatar split clause 44,77]
93. ~v1_finset_1(k10_xtuple_0(X0)) | v1_finset_1(k1_relat_1(X0)) | ~v1_finset_1(k9_xtuple_0(X0)) | ~v1_relat_1(X0) [superposition 60,50] {a:1,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
95. v1_finset_1(k1_relat_1(sK0)) | ~v1_finset_1(k9_xtuple_0(sK0)) | ~v1_relat_1(sK0) <- (2) [resolution 93,69] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
100. v1_finset_1(k1_relat_1(sK0)) | ~v1_relat_1(sK0) <- (2, 3) [subsumption resolution 95,74] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
101. v1_finset_1(k1_relat_1(sK0)) <- (2, 3, 4) [subsumption resolution 100,79] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
103. 5 <=> v1_finset_1(k1_relat_1(sK0)) [avatar definition]
105. v1_finset_1(k1_relat_1(sK0)) <- (5) [avatar component clause 103] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
106. 5 | ~2 | ~3 | ~4 [avatar split clause 101,77,72,67,103]
109. v1_finset_1(sK0) | ~v1_relat_1(sK0) <- (5) [resolution 105,51] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
111. ~v1_relat_1(sK0) <- (~1, 5) [subsumption resolution 109,64] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
112. $false <- (~1, 4, 5) [subsumption resolution 111,79] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
113. 1 | ~4 | ~5 [avatar contradiction clause 112]
114. $false [avatar sat refutation 65,70,75,80,106,113] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:34,thDist:-34}
% SZS output end Proof for orders_1__t87_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 1
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 48
% Active clauses: 21
% Passive clauses: 23
% Final active clauses: 21
% Final passive clauses: 1
% Main loop iterations started: 22
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Fw subsumption resolutions: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 7
% Unit resulting resolution: 1
% Forward superposition: 6
% Backward superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

