i: [1,54,0,2,1,1,0,0] t124_funct_4
a: [3,92,0,2,1,54]
s: 92
i: [1,53,0,2,1,1,0,0] t124_funct_4
a: [3,97,0,2,1,53]
s: 97
i: [1,52,0,2,1,1,0,0] t124_funct_4
a: [3,102,0,2,1,52]
s: 102
i: [1,51,0,2,1,1,0,0] t124_funct_4
a: [3,107,0,2,1,51]
s: 107
i: [1,50,0,2,1,1,0,0] t124_funct_4
a: [3,112,0,2,1,50]
s: 112
i: [1,49,0,2,1,1,0,0] t124_funct_4
a: [3,117,0,2,1,49]
s: 117
i: [1,55,0,3,1,1,0,0] t124_funct_4
a: [3,87,0,3,1,55]
s: 87
i: [1,57,0,5,1,1,0,0] t124_funct_4
a: [3,77,0,5,1,57]
s: 77
i: [1,56,0,5,1,1,0,0] t124_funct_4
a: [3,82,0,5,1,56]
s: 82
i: [1,73,0,2,1,0,0,1] rc1_funct_1
a: [3,129,0,2,1,73]
s: 129
i: [1,72,0,2,1,0,0,1] rc1_funct_1
a: [3,134,0,2,1,72]
s: 134
i: [1,66,0,3,1,0,0,1] reflexivity_r1_tarski
s: 66
i: [1,65,0,4,1,0,0,255] existence_m1_subset_1
s: 65
i: [1,67,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 67
i: [1,62,0,7,2,0,0,255] t3_subset
s: 62
i: [1,63,0,7,2,0,0,255] t3_subset
s: 63
d: [2,136,1,5,1,1,40,67,82]
a: [3,140,1,5,1,136]
s: 140
i: [1,58,0,8,3,0,0,255] cc2_relat_1
s: 58
d: [2,142,1,5,1,0,40,62,65]
s: 142
i: [1,64,0,9,3,0,0,1] t1_xboole_1
s: 64
i: [1,68,0,13,5,0,0,255] idempotence_k1_funct_4
a: [3,124,0,9,3,68]
s: 124
d: [2,147,1,6,2,1,40,64,87]
s: 147
i: [1,59,0,10,4,0,0,255] cc3_funct_1
s: 59
d: [2,151,2,5,1,1,40,147,142]
a: [3,155,2,5,1,151]
s: 155
i: [1,70,0,12,5,0,0,1] dt_k1_funct_4
s: 70
i: [1,71,0,12,5,0,0,1] dt_k1_funct_4
s: 71
d: [2,159,3,8,2,1,40,155,64]
s: 159
i: [1,69,0,17,5,0,0,255] rd3_funct_4
s: 69
d: [2,165,4,7,1,1,40,159,142]
a: [3,169,4,7,1,165]
s: 169
i: [1,60,0,18,6,0,0,1] t32_funct_4
s: 60
i: [1,61,0,22,8,0,0,1] t123_funct_4
s: 61
d: [2,180,1,13,5,1,40,60,82]
d: [2,182,1,11,4,2,21,180,92]
d: [2,183,1,9,3,3,21,182,97]
d: [2,184,1,7,2,4,21,183,112]
d: [2,185,1,5,1,5,21,184,117]
a: [3,189,1,5,1,185]
s: 189
d: [2,145,1,6,2,0,40,58,65]
s: 145
d: [2,181,2,13,5,1,40,60,140]
d: [2,191,2,11,4,2,21,181,112]
d: [2,192,2,9,3,3,21,191,117]
d: [2,193,2,7,2,4,21,192,92]
d: [2,194,2,5,1,5,21,193,97]
a: [3,198,2,5,1,194]
s: 198
d: [2,146,1,7,3,0,40,58,63]
s: 146
d: [2,157,1,8,3,0,40,59,65]
s: 157
d: [2,221,3,6,2,1,40,146,155]
d: [2,223,3,4,1,2,21,221,107]
a: [3,227,3,4,1,223]
s: 227
d: [2,158,1,9,4,0,40,59,63]
s: 158
d: [2,241,3,8,3,1,40,158,155]
d: [2,243,3,6,2,2,21,241,102]
d: [2,244,3,4,1,3,21,243,107]
a: [3,248,3,4,1,244]
s: 248
d: [2,203,1,24,10,1,44,61,124]
d: [2,206,1,16,6,1,18,203]
s: 206
d: [2,257,1,14,5,1,21,206,146]
d: [2,258,1,12,4,1,21,257,158]
s: 258
d: [2,222,5,8,2,1,40,146,169]
d: [2,229,5,6,1,2,21,222,107]
a: [3,233,5,6,1,229]
s: 233
d: [2,201,1,24,10,1,44,61,124]
d: [2,208,1,16,6,1,18,201]
s: 208
d: [2,242,5,10,3,1,40,158,169]
d: [2,250,5,8,2,2,21,242,102]
d: [2,251,5,6,1,3,21,250,107]
a: [3,255,5,6,1,251]
s: 255
d: [2,274,1,14,5,1,21,208,146]
d: [2,275,1,12,4,1,21,274,158]
s: 275
d: [2,202,1,36,12,0,44,61,69]
d: [2,207,1,32,10,0,18,202]
d: [2,209,1,28,9,0,21,207,71]
d: [2,210,1,24,8,0,21,209,70]
s: 210
d: [2,213,2,8,2,1,40,189,64]
s: 213
d: [2,204,1,36,12,0,44,61,69]
d: [2,205,1,32,10,0,18,204]
d: [2,211,1,28,9,0,21,205,71]
d: [2,212,1,24,8,0,21,211,70]
s: 212
d: [2,293,3,13,4,1,40,213,158]
a: [3,299,3,4,1,293]
s: 299
d: [2,323,4,0,0,5,52,117,112,97,92,299,70]
e: 323
d: [2,324,4,8,4,1,40,299,70]
d: [2,326,4,6,3,2,21,324,92]
d: [2,327,4,4,2,3,21,326,97]
d: [2,328,4,2,1,4,21,327,112]
d: [2,329,4,0,0,5,21,328,117]
e: 329
r: 299
d: [2,200,1,25,9,0,40,61,64]
s: 200
d: [2,331,1,23,8,0,21,200,146]
d: [2,332,1,21,7,0,21,331,158]
s: 332
d: [2,333,2,0,0,7,52,97,92,107,102,87,77,189,332]
e: 333
f: 333,185,49,50,51,52,53,54,55,56,57
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t124_funct_4
% SZS output start Proof for funct_4__t124_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X2)) & r1_tarski(X0,X1)) => r1_tarski(X0,k1_funct_4(X1,X2)))))) [input t124_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X2)) & r1_tarski(X0,X1)) => r1_tarski(X0,k1_funct_4(X1,X2)))))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
4. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_funct_1(X1))) [input cc3_funct_1]
15. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r1_tarski(X0,X1) => r1_tarski(k1_funct_4(X0,X2),k1_funct_4(X1,X2)))))) [input t123_funct_4]
16. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
17. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => r1_tarski(X0,k1_funct_4(X0,X1))))) [input t32_funct_4]
18. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
19. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 18]
20. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 16]
22. ? [X0] : (? [X1] : (? [X2] : ((~r1_tarski(X0,k1_funct_4(X1,X2)) & (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X2)) & r1_tarski(X0,X1))) & (v1_funct_1(X2) & v1_relat_1(X2))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
23. ? [X0] : (? [X1] : (? [X2] : (~r1_tarski(X0,k1_funct_4(X1,X2)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X2)) & r1_tarski(X0,X1) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 22]
24. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
25. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 4]
26. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 25]
27. ! [X0] : (! [X1] : ((r1_tarski(X0,k1_funct_4(X0,X1)) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 17]
28. ! [X0] : (! [X1] : (r1_tarski(X0,k1_funct_4(X0,X1)) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 27]
29. ! [X0] : (! [X1] : (! [X2] : ((r1_tarski(k1_funct_4(X0,X2),k1_funct_4(X1,X2)) | ~r1_tarski(X0,X1)) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 15]
30. ! [X0] : (! [X1] : (! [X2] : (r1_tarski(k1_funct_4(X0,X2),k1_funct_4(X1,X2)) | ~r1_tarski(X0,X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 29]
31. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 20]
32. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 31]
40. ? [X0] : (? [X1] : (? [X2] : (~r1_tarski(X0,k1_funct_4(X1,X2)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X2)) & r1_tarski(X0,X1) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (~r1_tarski(sK0,k1_funct_4(X1,X2)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X2)) & r1_tarski(sK0,X1) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
41. ? [X1] : (? [X2] : (~r1_tarski(sK0,k1_funct_4(X1,X2)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X2)) & r1_tarski(sK0,X1) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) => (? [X2] : (~r1_tarski(sK0,k1_funct_4(sK1,X2)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X2)) & r1_tarski(sK0,sK1) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
42. ? [X2] : (~r1_tarski(sK0,k1_funct_4(sK1,X2)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X2)) & r1_tarski(sK0,sK1) & v1_funct_1(X2) & v1_relat_1(X2)) => (~r1_tarski(sK0,k1_funct_4(sK1,sK2)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK2)) & r1_tarski(sK0,sK1) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
43. ((~r1_tarski(sK0,k1_funct_4(sK1,sK2)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK2)) & r1_tarski(sK0,sK1) & v1_funct_1(sK2) & v1_relat_1(sK2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 23,42,41,40]
44. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 19]
49. v1_relat_1(sK0) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
50. v1_funct_1(sK0) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
51. v1_relat_1(sK1) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
52. v1_funct_1(sK1) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
53. v1_relat_1(sK2) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
54. v1_funct_1(sK2) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
55. r1_tarski(sK0,sK1) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
56. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK2)) [cnf transformation 43] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
57. ~r1_tarski(sK0,k1_funct_4(sK1,sK2)) [cnf transformation 43] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
58. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
59. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_funct_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
60. ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | r1_tarski(X0,k1_funct_4(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. r1_tarski(k1_funct_4(X0,X2),k1_funct_4(X1,X2)) | ~r1_tarski(X0,X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 44] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
64. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 32] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
75. 1 <=> r1_tarski(sK0,k1_funct_4(sK1,sK2)) [avatar definition]
77. ~r1_tarski(sK0,k1_funct_4(sK1,sK2)) <- (~1) [avatar component clause 75] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
78. ~1 [avatar split clause 57,75]
80. 2 <=> r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK2)) [avatar definition]
82. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK2)) <- (2) [avatar component clause 80] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
83. 2 [avatar split clause 56,80]
85. 3 <=> r1_tarski(sK0,sK1) [avatar definition]
87. r1_tarski(sK0,sK1) <- (3) [avatar component clause 85] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
88. 3 [avatar split clause 55,85]
90. 4 <=> v1_funct_1(sK2) [avatar definition]
92. v1_funct_1(sK2) <- (4) [avatar component clause 90] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
93. 4 [avatar split clause 54,90]
95. 5 <=> v1_relat_1(sK2) [avatar definition]
97. v1_relat_1(sK2) <- (5) [avatar component clause 95] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
98. 5 [avatar split clause 53,95]
100. 6 <=> v1_funct_1(sK1) [avatar definition]
102. v1_funct_1(sK1) <- (6) [avatar component clause 100] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
103. 6 [avatar split clause 52,100]
105. 7 <=> v1_relat_1(sK1) [avatar definition]
107. v1_relat_1(sK1) <- (7) [avatar component clause 105] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
108. 7 [avatar split clause 51,105]
110. 8 <=> v1_funct_1(sK0) [avatar definition]
112. v1_funct_1(sK0) <- (8) [avatar component clause 110] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
113. 8 [avatar split clause 50,110]
115. 9 <=> v1_relat_1(sK0) [avatar definition]
117. v1_relat_1(sK0) <- (9) [avatar component clause 115] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
118. 9 [avatar split clause 49,115]
146. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 58,63] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
158. ~r1_tarski(X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | v1_funct_1(X1) [resolution 59,63] {a:1,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
180. r1_tarski(sK0,k1_funct_4(sK0,sK2)) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 60,82] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
182. r1_tarski(sK0,k1_funct_4(sK0,sK2)) | ~v1_relat_1(sK2) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 4) [subsumption resolution 180,92] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
183. r1_tarski(sK0,k1_funct_4(sK0,sK2)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 4, 5) [subsumption resolution 182,97] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
184. r1_tarski(sK0,k1_funct_4(sK0,sK2)) | ~v1_relat_1(sK0) <- (2, 4, 5, 8) [subsumption resolution 183,112] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
185. r1_tarski(sK0,k1_funct_4(sK0,sK2)) <- (2, 4, 5, 8, 9) [subsumption resolution 184,117] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
187. 17 <=> r1_tarski(sK0,k1_funct_4(sK0,sK2)) [avatar definition]
189. r1_tarski(sK0,k1_funct_4(sK0,sK2)) <- (17) [avatar component clause 187] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
190. 17 | ~2 | ~4 | ~5 | ~8 | ~9 [avatar split clause 185,115,110,95,90,80,187]
200. ~r1_tarski(X0,X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | r1_tarski(X3,k1_funct_4(X1,X2)) | ~r1_tarski(X3,k1_funct_4(X0,X2)) [resolution 61,64] {a:1,w:25,wCS:10000,sine:1,thAx:0,allAx:2,thDist:-2}
331. ~r1_tarski(X0,X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | r1_tarski(X3,k1_funct_4(X1,X2)) | ~r1_tarski(X3,k1_funct_4(X0,X2)) [subsumption resolution 200,146] {a:1,w:23,wCS:9200,sine:1,thAx:0,allAx:2,thDist:-2}
332. ~r1_tarski(X3,k1_funct_4(X0,X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | r1_tarski(X3,k1_funct_4(X1,X2)) | ~r1_tarski(X0,X1) [subsumption resolution 331,158] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
333. $false <- (~1, 3, 4, 5, 6, 7, 17) [unit resulting resolution 97,92,107,102,87,77,189,332] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:31,thDist:-31}
351. 1 | ~3 | ~4 | ~5 | ~6 | ~7 | ~17 [avatar contradiction clause 333]
369. $false [avatar sat refutation 78,83,88,93,98,103,108,113,118,190,351] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:258,thDist:-258}
% SZS output end Proof for funct_4__t124_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 25
% Generated clauses: 219
% Active clauses: 48
% Passive clauses: 66
% Final active clauses: 47
% Final passive clauses: 15
% Main loop iterations started: 51
% 
% >>> Simplifying Inferences
% Duplicate literals: 73
% Fw subsumption resolutions: 49
% 
% >>> Deletion Inferences
% Simple tautologies: 16
% Forward subsumptions: 25
% 
% >>> Generating Inferences
% Binary resolution: 66
% Unit resulting resolution: 2
% Forward superposition: 21
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 28
% SAT solver unit clauses: 11
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

