i: [1,35,0,2,1,1,0,0] t30_xxreal_0
a: [3,59,0,2,1,35]
s: 59
i: [1,34,0,2,1,1,0,0] t30_xxreal_0
a: [3,64,0,2,1,34]
s: 64
i: [1,33,0,2,1,1,0,0] t30_xxreal_0
a: [3,69,0,2,1,33]
s: 69
i: [1,37,0,3,1,1,0,0] t30_xxreal_0
a: [3,49,0,3,1,37]
s: 49
i: [1,36,0,5,1,1,0,0] t30_xxreal_0
a: [3,54,0,5,1,36]
s: 54
i: [1,45,0,2,1,0,0,1] rc1_xxreal_0
a: [3,85,0,2,1,45]
s: 85
i: [1,40,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,76,0,5,2,40]
s: 76
i: [1,42,0,9,3,0,0,255] idempotence_k4_xxreal_0
a: [3,80,0,7,2,42]
s: 80
i: [1,41,0,8,3,0,0,1] fc6_xxreal_0
s: 41
i: [1,38,0,9,3,0,0,1] t25_xxreal_0
s: 38
i: [1,44,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 44
d: [2,91,1,7,3,1,40,44,49]
d: [2,94,1,5,2,2,21,91,69]
d: [2,95,1,3,1,3,21,94,59]
a: [3,99,1,3,1,95]
s: 99
i: [1,43,0,11,3,0,0,255] commutativity_k4_xxreal_0
s: 43
d: [2,103,1,13,5,0,44,38,43]
d: [2,106,1,9,3,0,18,103]
s: 106
i: [1,39,0,15,6,0,0,1] t2_xxreal_0
s: 39
d: [2,115,1,16,5,1,40,39,54]
d: [2,125,1,14,4,2,21,115,59]
a: [3,129,1,4,1,125]
s: 129
d: [2,138,2,0,0,3,52,69,64,129,41]
e: 138
d: [2,139,2,4,2,1,40,129,41]
d: [2,141,2,2,1,2,21,139,64]
d: [2,142,2,0,0,3,21,141,69]
e: 142
r: 129
a: [3,128,1,4,1,125]
s: 128
a: [3,132,1,10,3,125]
s: 132
d: [2,144,2,9,4,1,40,132,38]
d: [2,151,2,7,3,1,18,144]
d: [2,152,2,5,2,2,21,151,69]
d: [2,153,2,2,1,3,21,152,49]
d: [2,154,2,0,0,4,21,153,64]
e: 154
f: 154,142,125,33,34,35,36,37
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_0__t30_xxreal_0
% SZS output start Proof for xxreal_0__t30_xxreal_0
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r1_xxreal_0(k4_xxreal_0(X0,X1),X2) => r1_xxreal_0(X0,X2))))) [input t30_xxreal_0]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r1_xxreal_0(k4_xxreal_0(X0,X1),X2) => r1_xxreal_0(X0,X2))))) [negated conjecture 1]
6. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v1_xxreal_0(k4_xxreal_0(X0,X1))) [input fc6_xxreal_0]
10. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => r1_xxreal_0(X0,k4_xxreal_0(X0,X1)))) [input t25_xxreal_0]
11. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input t2_xxreal_0]
12. ? [X0] : (? [X1] : (? [X2] : ((~r1_xxreal_0(X0,X2) & r1_xxreal_0(k4_xxreal_0(X0,X1),X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
13. ? [X0] : (? [X1] : (? [X2] : (~r1_xxreal_0(X0,X2) & r1_xxreal_0(k4_xxreal_0(X0,X1),X2) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 12]
14. ! [X0] : (! [X1] : (r1_xxreal_0(X0,k4_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 10]
15. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 11]
16. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 15]
19. ! [X0,X1] : (v1_xxreal_0(k4_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 6]
20. ! [X0,X1] : (v1_xxreal_0(k4_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 19]
27. ? [X0] : (? [X1] : (? [X2] : (~r1_xxreal_0(X0,X2) & r1_xxreal_0(k4_xxreal_0(X0,X1),X2) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (~r1_xxreal_0(sK0,X2) & r1_xxreal_0(k4_xxreal_0(sK0,X1),X2) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
28. ? [X1] : (? [X2] : (~r1_xxreal_0(sK0,X2) & r1_xxreal_0(k4_xxreal_0(sK0,X1),X2) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (~r1_xxreal_0(sK0,X2) & r1_xxreal_0(k4_xxreal_0(sK0,sK1),X2) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
29. ? [X2] : (~r1_xxreal_0(sK0,X2) & r1_xxreal_0(k4_xxreal_0(sK0,sK1),X2) & v1_xxreal_0(X2)) => (~r1_xxreal_0(sK0,sK2) & r1_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) & v1_xxreal_0(sK2)) [choice axiom]
30. ((~r1_xxreal_0(sK0,sK2) & r1_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 13,29,28,27]
33. v1_xxreal_0(sK0) [cnf transformation 30] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
34. v1_xxreal_0(sK1) [cnf transformation 30] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
35. v1_xxreal_0(sK2) [cnf transformation 30] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
36. r1_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) [cnf transformation 30] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
37. ~r1_xxreal_0(sK0,sK2) [cnf transformation 30] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
38. r1_xxreal_0(X0,k4_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 14] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
39. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 16] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
41. v1_xxreal_0(k4_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 20] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. 1 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
49. ~r1_xxreal_0(sK0,sK2) <- (~1) [avatar component clause 47] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
50. ~1 [avatar split clause 37,47]
52. 2 <=> r1_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) [avatar definition]
54. r1_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) <- (2) [avatar component clause 52] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
55. 2 [avatar split clause 36,52]
57. 3 <=> v1_xxreal_0(sK2) [avatar definition]
59. v1_xxreal_0(sK2) <- (3) [avatar component clause 57] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
60. 3 [avatar split clause 35,57]
62. 4 <=> v1_xxreal_0(sK1) [avatar definition]
64. v1_xxreal_0(sK1) <- (4) [avatar component clause 62] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
65. 4 [avatar split clause 34,62]
67. 5 <=> v1_xxreal_0(sK0) [avatar definition]
69. v1_xxreal_0(sK0) <- (5) [avatar component clause 67] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
70. 5 [avatar split clause 33,67]
115. r1_xxreal_0(X0,sK2) | ~r1_xxreal_0(X0,k4_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(k4_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(X0) <- (2) [resolution 39,54] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
125. r1_xxreal_0(X0,sK2) | ~r1_xxreal_0(X0,k4_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(k4_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(X0) <- (2, 3) [subsumption resolution 115,59] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
127. 11 <=> v1_xxreal_0(k4_xxreal_0(sK0,sK1)) [avatar definition]
129. ~v1_xxreal_0(k4_xxreal_0(sK0,sK1)) <- (~11) [avatar component clause 127] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
131. 12 <=> ! [X0] : (r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X0) | ~r1_xxreal_0(X0,k4_xxreal_0(sK0,sK1))) [avatar definition]
132. ~r1_xxreal_0(X0,k4_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(X0) | r1_xxreal_0(X0,sK2) <- (12) [avatar component clause 131] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
133. ~11 | 12 | ~2 | ~3 [avatar split clause 125,57,52,131,127]
139. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~11) [resolution 129,41] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
141. ~v1_xxreal_0(sK0) <- (4, ~11) [subsumption resolution 139,64] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
142. $false <- (4, 5, ~11) [subsumption resolution 141,69] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
143. ~4 | ~5 | 11 [avatar contradiction clause 142]
144. ~v1_xxreal_0(sK0) | r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (12) [resolution 132,38] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
151. ~v1_xxreal_0(sK0) | r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK1) <- (12) [duplicate literal removal 144] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
152. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK1) <- (5, 12) [subsumption resolution 151,69] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
153. ~v1_xxreal_0(sK1) <- (~1, 5, 12) [subsumption resolution 152,49] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
154. $false <- (~1, 4, 5, 12) [subsumption resolution 153,64] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
155. 1 | ~4 | ~5 | ~12 [avatar contradiction clause 154]
159. $false [avatar sat refutation 50,55,60,65,70,133,143,155] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:78,thDist:-78}
% SZS output end Proof for xxreal_0__t30_xxreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 81
% Active clauses: 18
% Passive clauses: 23
% Final active clauses: 17
% Final passive clauses: 5
% Main loop iterations started: 18
% 
% >>> Simplifying Inferences
% Duplicate literals: 29
% Fw subsumption resolutions: 15
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 12
% Unit resulting resolution: 1
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 4
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 6
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward demodulation: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

