i: [1,51,0,2,1,1,0,0] t55_rewrite1
a: [3,89,0,2,1,51]
s: 89
i: [1,50,0,2,1,1,0,0] t55_rewrite1
a: [3,94,0,2,1,50]
s: 94
i: [1,49,0,2,1,1,0,0] t55_rewrite1
a: [3,99,0,2,1,49]
s: 99
i: [1,52,0,4,1,1,0,0] t55_rewrite1
a: [3,84,0,4,1,52]
s: 84
i: [1,53,0,7,1,1,0,0] t55_rewrite1
a: [3,79,0,7,1,53]
s: 79
i: [1,66,0,6,3,0,0,255] cc7_rewrite1
s: 66
i: [1,68,0,6,3,0,0,255] cc10_rewrite1
s: 68
i: [1,73,0,6,3,0,0,255] cc6_rewrite1
s: 73
i: [1,75,0,6,3,0,0,255] cc9_rewrite1
s: 75
i: [1,55,0,8,3,0,0,1] d19_rewrite1
s: 55
i: [1,56,0,8,3,0,0,1] d19_rewrite1
s: 56
i: [1,70,0,8,4,0,0,1] cc11_rewrite1
s: 70
i: [1,58,0,9,3,0,0,1] d19_rewrite1
s: 58
d: [2,101,1,6,3,1,40,70,89]
d: [2,102,1,4,2,2,21,101,94]
d: [2,103,1,2,1,3,21,102,99]
a: [3,107,1,2,1,103]
s: 107
i: [1,61,0,9,3,0,0,2] d6_rewrite1
s: 61
i: [1,57,0,10,3,0,0,1] d19_rewrite1
s: 57
d: [2,109,2,4,2,1,40,107,75]
d: [2,111,2,2,1,2,21,109,99]
a: [3,115,2,2,1,111]
s: 115
i: [1,59,0,10,3,0,0,3] t25_rewrite1
s: 59
d: [2,110,2,4,2,1,40,107,73]
d: [2,117,2,2,1,2,21,110,99]
a: [3,121,2,2,1,117]
s: 121
i: [1,60,0,10,3,0,0,3] t25_rewrite1
s: 60
i: [1,62,0,10,3,0,0,2] d6_rewrite1
s: 62
i: [1,71,0,12,4,0,0,1] t54_rewrite1
s: 71
i: [1,63,0,13,4,0,0,2] d6_rewrite1
s: 63
d: [2,126,1,13,5,0,40,71,61]
d: [2,127,1,11,4,0,18,126]
s: 127
i: [1,64,0,14,4,0,0,1] t30_rewrite1
s: 64
i: [1,54,0,17,6,0,0,1] d19_rewrite1
s: 54
d: [2,133,1,13,5,1,40,54,84]
d: [2,136,1,11,4,2,21,133,89]
d: [2,137,1,9,3,3,21,136,99]
a: [3,141,1,3,1,137]
s: 141
d: [2,129,1,10,3,1,40,64,84]
d: [2,132,1,8,2,2,21,129,99]
s: 132
d: [2,125,1,14,5,0,40,71,62]
d: [2,128,1,12,4,0,18,125]
s: 128
d: [2,130,1,16,5,0,40,64,57]
d: [2,131,1,14,4,0,18,130]
s: 131
d: [2,153,2,14,5,0,40,128,59]
d: [2,154,2,12,4,0,18,153]
s: 154
d: [2,152,2,14,5,0,40,128,60]
d: [2,155,2,12,4,0,18,152]
s: 155
d: [2,165,3,12,4,2,40,155,132]
d: [2,170,3,10,3,2,21,165,99]
d: [2,171,3,8,2,3,21,170,89]
d: [2,172,3,6,1,4,21,171,94]
a: [3,176,3,6,1,172]
s: 176
a: [3,145,1,3,1,137]
s: 145
d: [2,157,3,23,8,0,40,154,54]
d: [2,160,3,19,6,0,18,157]
d: [2,161,3,14,5,0,21,160,127]
s: 161
d: [2,162,3,23,8,0,40,155,54]
d: [2,168,3,19,6,0,18,162]
d: [2,169,3,14,5,0,21,168,127]
s: 169
d: [2,179,4,12,3,1,40,176,64]
d: [2,191,4,10,2,2,21,179,99]
s: 191
d: [2,158,3,18,6,0,40,154,64]
d: [2,159,3,16,5,0,18,158]
s: 159
d: [2,205,5,14,4,2,40,191,155]
d: [2,207,5,12,3,2,21,205,99]
d: [2,208,5,10,2,3,21,207,89]
d: [2,209,5,8,1,4,21,208,94]
a: [3,213,5,8,1,209]
s: 213
d: [2,163,3,18,6,0,40,155,64]
d: [2,167,3,16,5,0,18,163]
s: 167
d: [2,228,4,16,7,0,40,167,154]
d: [2,235,4,10,4,0,18,228]
s: 235
d: [2,224,6,14,3,1,40,213,64]
d: [2,227,6,12,2,2,21,224,99]
s: 227
d: [2,198,4,19,8,0,44,71,161]
d: [2,199,4,13,5,0,18,198]
s: 199
d: [2,249,7,16,4,2,40,227,155]
d: [2,252,7,14,3,2,21,249,99]
d: [2,253,7,12,2,3,21,252,89]
d: [2,254,7,10,1,4,21,253,94]
a: [3,258,7,10,1,254]
s: 258
d: [2,196,4,19,8,0,44,128,161]
d: [2,201,4,13,5,0,18,196]
s: 201
d: [2,195,4,19,8,0,44,154,161]
d: [2,202,4,13,5,0,18,195]
s: 202
d: [2,265,8,16,3,1,40,258,64]
d: [2,269,8,14,2,2,21,265,99]
s: 269
d: [2,194,4,19,8,0,44,155,161]
d: [2,203,4,13,5,0,18,194]
s: 203
d: [2,279,9,18,4,2,40,269,155]
d: [2,283,9,16,3,2,21,279,99]
d: [2,284,9,14,2,3,21,283,89]
d: [2,285,9,12,1,4,21,284,94]
a: [3,289,9,12,1,285]
s: 289
d: [2,230,4,20,7,0,40,167,155]
d: [2,233,4,14,4,0,18,230]
s: 233
d: [2,215,4,27,9,0,40,159,54]
d: [2,221,4,23,7,0,18,215]
d: [2,222,4,18,6,0,21,221,127]
s: 222
d: [2,325,5,18,6,1,44,79,222]
d: [2,364,5,16,5,2,21,325,94]
d: [2,365,5,14,4,3,21,364,99]
d: [2,366,5,12,3,4,21,365,89]
s: 366
d: [2,178,4,17,5,1,40,176,54]
d: [2,180,4,15,4,2,21,178,89]
d: [2,181,4,13,3,3,21,180,99]
a: [3,185,4,5,1,181]
s: 185
d: [2,411,5,0,0,4,52,89,99,94,185,127]
e: 411
d: [2,412,5,6,3,1,40,185,127]
d: [2,418,5,4,2,2,21,412,94]
d: [2,419,5,2,1,3,21,418,99]
d: [2,420,5,0,0,4,21,419,89]
e: 420
d: [2,405,6,11,2,4,47,366]
d: [2,409,6,5,1,5,21,405,176]
f: 420,409,172,49,50,51,52,53
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rewrite1__t55_rewrite1
% SZS output start Proof for rewrite1__t55_rewrite1
1. ! [X0] : ((v4_rewrite1(X0) & v2_rewrite1(X0) & v1_relat_1(X0)) => ! [X1] : ! [X2] : (r2_rewrite1(X0,X1,X2) => k2_rewrite1(X0,X1) = k2_rewrite1(X0,X2))) [input t55_rewrite1]
2. ~! [X0] : ((v4_rewrite1(X0) & v2_rewrite1(X0) & v1_relat_1(X0)) => ! [X1] : ! [X2] : (r2_rewrite1(X0,X1,X2) => k2_rewrite1(X0,X1) = k2_rewrite1(X0,X2))) [negated conjecture 1]
8. ! [X0] : (v1_relat_1(X0) => (v4_rewrite1(X0) <=> ! [X1] : ! [X2] : ((r2_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2) & r3_rewrite1(X0,X1)) => X1 = X2))) [input d19_rewrite1]
9. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r4_rewrite1(X0,X1,X2) <=> (r1_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2)))) [input d6_rewrite1]
11. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r1_rewrite1(X0,X1,X2) => (r2_rewrite1(X0,X2,X1) & r2_rewrite1(X0,X1,X2)))) [input t25_rewrite1]
12. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : ((r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X3))) [input t30_rewrite1]
13. ! [X0] : ((v4_rewrite1(X0) & v2_rewrite1(X0) & v1_relat_1(X0)) => ! [X1] : r4_rewrite1(X0,X1,k2_rewrite1(X0,X1))) [input t54_rewrite1]
14. ~! [X0] : ((v4_rewrite1(X0) & v2_rewrite1(X0) & v1_relat_1(X0)) => ! [X1,X2] : (r2_rewrite1(X0,X1,X2) => k2_rewrite1(X0,X1) = k2_rewrite1(X0,X2))) [flattening 2]
15. ! [X0] : (v1_relat_1(X0) => (v4_rewrite1(X0) <=> ! [X1,X2] : ((r2_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2) & r3_rewrite1(X0,X1)) => X1 = X2))) [flattening 8]
16. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r1_rewrite1(X0,X1,X2) => (r2_rewrite1(X0,X2,X1) & r2_rewrite1(X0,X1,X2)))) [flattening 11]
17. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r4_rewrite1(X0,X1,X2) <=> (r1_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2)))) [flattening 9]
18. ! [X0] : (v1_relat_1(X0) => ! [X1,X2,X3] : ((r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X3))) [flattening 12]
20. ? [X0] : (? [X1,X2] : (k2_rewrite1(X0,X1) != k2_rewrite1(X0,X2) & r2_rewrite1(X0,X1,X2)) & (v4_rewrite1(X0) & v2_rewrite1(X0) & v1_relat_1(X0))) [ennf transformation 14]
21. ? [X0] : (? [X1,X2] : (k2_rewrite1(X0,X1) != k2_rewrite1(X0,X2) & r2_rewrite1(X0,X1,X2)) & v4_rewrite1(X0) & v2_rewrite1(X0) & v1_relat_1(X0)) [flattening 20]
22. ! [X0] : ((v4_rewrite1(X0) <=> ! [X1,X2] : (X1 = X2 | (~r2_rewrite1(X0,X1,X2) | ~r3_rewrite1(X0,X2) | ~r3_rewrite1(X0,X1)))) | ~v1_relat_1(X0)) [ennf transformation 15]
23. ! [X0] : ((v4_rewrite1(X0) <=> ! [X1,X2] : (X1 = X2 | ~r2_rewrite1(X0,X1,X2) | ~r3_rewrite1(X0,X2) | ~r3_rewrite1(X0,X1))) | ~v1_relat_1(X0)) [flattening 22]
24. ! [X0] : (! [X1,X2] : ((r2_rewrite1(X0,X2,X1) & r2_rewrite1(X0,X1,X2)) | ~r1_rewrite1(X0,X1,X2)) | ~v1_relat_1(X0)) [ennf transformation 16]
25. ! [X0] : (! [X1,X2] : (r4_rewrite1(X0,X1,X2) <=> (r1_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2))) | ~v1_relat_1(X0)) [ennf transformation 17]
26. ! [X0] : (! [X1,X2,X3] : (r2_rewrite1(X0,X1,X3) | (~r2_rewrite1(X0,X2,X3) | ~r2_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [ennf transformation 18]
27. ! [X0] : (! [X1,X2,X3] : (r2_rewrite1(X0,X1,X3) | ~r2_rewrite1(X0,X2,X3) | ~r2_rewrite1(X0,X1,X2)) | ~v1_relat_1(X0)) [flattening 26]
34. ! [X0] : (! [X1] : r4_rewrite1(X0,X1,k2_rewrite1(X0,X1)) | (~v4_rewrite1(X0) | ~v2_rewrite1(X0) | ~v1_relat_1(X0))) [ennf transformation 13]
35. ! [X0] : (! [X1] : r4_rewrite1(X0,X1,k2_rewrite1(X0,X1)) | ~v4_rewrite1(X0) | ~v2_rewrite1(X0) | ~v1_relat_1(X0)) [flattening 34]
40. ? [X0] : (? [X1,X2] : (k2_rewrite1(X0,X1) != k2_rewrite1(X0,X2) & r2_rewrite1(X0,X1,X2)) & v4_rewrite1(X0) & v2_rewrite1(X0) & v1_relat_1(X0)) => (? [X2,X1] : (k2_rewrite1(sK0,X1) != k2_rewrite1(sK0,X2) & r2_rewrite1(sK0,X1,X2)) & v4_rewrite1(sK0) & v2_rewrite1(sK0) & v1_relat_1(sK0)) [choice axiom]
41. ? [X2,X1] : (k2_rewrite1(sK0,X1) != k2_rewrite1(sK0,X2) & r2_rewrite1(sK0,X1,X2)) => (k2_rewrite1(sK0,sK1) != k2_rewrite1(sK0,sK2) & r2_rewrite1(sK0,sK1,sK2)) [choice axiom]
42. (k2_rewrite1(sK0,sK1) != k2_rewrite1(sK0,sK2) & r2_rewrite1(sK0,sK1,sK2)) & v4_rewrite1(sK0) & v2_rewrite1(sK0) & v1_relat_1(sK0) [skolemisation 21,41,40]
43. ! [X0] : (((v4_rewrite1(X0) | ? [X1,X2] : (X1 != X2 & r2_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2) & r3_rewrite1(X0,X1))) & (! [X1,X2] : (X1 = X2 | ~r2_rewrite1(X0,X1,X2) | ~r3_rewrite1(X0,X2) | ~r3_rewrite1(X0,X1)) | ~v4_rewrite1(X0))) | ~v1_relat_1(X0)) [nnf transformation 23]
44. ! [X0] : (((v4_rewrite1(X0) | ? [X1,X2] : (X1 != X2 & r2_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2) & r3_rewrite1(X0,X1))) & (! [X3,X4] : (X3 = X4 | ~r2_rewrite1(X0,X3,X4) | ~r3_rewrite1(X0,X4) | ~r3_rewrite1(X0,X3)) | ~v4_rewrite1(X0))) | ~v1_relat_1(X0)) [rectify 43]
45. ! [X0] : (? [X1,X2] : (X1 != X2 & r2_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2) & r3_rewrite1(X0,X1)) => (sK3(X0) != sK4(X0) & r2_rewrite1(X0,sK3(X0),sK4(X0)) & r3_rewrite1(X0,sK4(X0)) & r3_rewrite1(X0,sK3(X0)))) [choice axiom]
46. ! [X0] : (((v4_rewrite1(X0) | (sK3(X0) != sK4(X0) & r2_rewrite1(X0,sK3(X0),sK4(X0)) & r3_rewrite1(X0,sK4(X0)) & r3_rewrite1(X0,sK3(X0)))) & (! [X3,X4] : (X3 = X4 | ~r2_rewrite1(X0,X3,X4) | ~r3_rewrite1(X0,X4) | ~r3_rewrite1(X0,X3)) | ~v4_rewrite1(X0))) | ~v1_relat_1(X0)) [skolemisation 44,45]
47. ! [X0] : (! [X1,X2] : ((r4_rewrite1(X0,X1,X2) | (~r1_rewrite1(X0,X1,X2) | ~r3_rewrite1(X0,X2))) & ((r1_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2)) | ~r4_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [nnf transformation 25]
48. ! [X0] : (! [X1,X2] : ((r4_rewrite1(X0,X1,X2) | ~r1_rewrite1(X0,X1,X2) | ~r3_rewrite1(X0,X2)) & ((r1_rewrite1(X0,X1,X2) & r3_rewrite1(X0,X2)) | ~r4_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [flattening 47]
49. v1_relat_1(sK0) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. v2_rewrite1(sK0) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. v4_rewrite1(sK0) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. r2_rewrite1(sK0,sK1,sK2) [cnf transformation 42] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
53. k2_rewrite1(sK0,sK1) != k2_rewrite1(sK0,sK2) [cnf transformation 42] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. ~r2_rewrite1(X0,X3,X4) | X3 = X4 | ~r3_rewrite1(X0,X4) | ~r3_rewrite1(X0,X3) | ~v4_rewrite1(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
59. ~r1_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:10,wCS:4000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
60. ~r1_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X2,X1) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:10,wCS:4000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
61. ~r4_rewrite1(X0,X1,X2) | r3_rewrite1(X0,X2) | ~v1_relat_1(X0) [cnf transformation 48] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
62. ~r4_rewrite1(X0,X1,X2) | r1_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 48] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
64. ~r2_rewrite1(X0,X2,X3) | r2_rewrite1(X0,X1,X3) | ~r2_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. r4_rewrite1(X0,X1,k2_rewrite1(X0,X1)) | ~v4_rewrite1(X0) | ~v2_rewrite1(X0) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
77. 1 <=> k2_rewrite1(sK0,sK1) = k2_rewrite1(sK0,sK2) [avatar definition]
79. k2_rewrite1(sK0,sK1) != k2_rewrite1(sK0,sK2) <- (~1) [avatar component clause 77] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
80. ~1 [avatar split clause 53,77]
82. 2 <=> r2_rewrite1(sK0,sK1,sK2) [avatar definition]
84. r2_rewrite1(sK0,sK1,sK2) <- (2) [avatar component clause 82] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
85. 2 [avatar split clause 52,82]
87. 3 <=> v4_rewrite1(sK0) [avatar definition]
89. v4_rewrite1(sK0) <- (3) [avatar component clause 87] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
90. 3 [avatar split clause 51,87]
92. 4 <=> v2_rewrite1(sK0) [avatar definition]
94. v2_rewrite1(sK0) <- (4) [avatar component clause 92] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. 4 [avatar split clause 50,92]
97. 5 <=> v1_relat_1(sK0) [avatar definition]
99. v1_relat_1(sK0) <- (5) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. 5 [avatar split clause 49,97]
125. ~v4_rewrite1(X0) | ~v2_rewrite1(X0) | ~v1_relat_1(X0) | r1_rewrite1(X0,X1,k2_rewrite1(X0,X1)) | ~v1_relat_1(X0) [resolution 71,62] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
126. ~v4_rewrite1(X2) | ~v2_rewrite1(X2) | ~v1_relat_1(X2) | r3_rewrite1(X2,k2_rewrite1(X2,X3)) | ~v1_relat_1(X2) [resolution 71,61] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
127. r3_rewrite1(X2,k2_rewrite1(X2,X3)) | ~v2_rewrite1(X2) | ~v1_relat_1(X2) | ~v4_rewrite1(X2) [duplicate literal removal 126] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
128. r1_rewrite1(X0,X1,k2_rewrite1(X0,X1)) | ~v2_rewrite1(X0) | ~v1_relat_1(X0) | ~v4_rewrite1(X0) [duplicate literal removal 125] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
129. r2_rewrite1(sK0,X0,sK2) | ~r2_rewrite1(sK0,X0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 64,84] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
132. ~r2_rewrite1(sK0,X0,sK1) | r2_rewrite1(sK0,X0,sK2) <- (2, 5) [subsumption resolution 129,99] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
152. ~v2_rewrite1(X2) | ~v1_relat_1(X2) | ~v4_rewrite1(X2) | r2_rewrite1(X2,k2_rewrite1(X2,X3),X3) | ~v1_relat_1(X2) [resolution 128,60] {a:2,w:14,wCS:5600,sine:1,thAx:0,allAx:3,thDist:-3}
153. ~v2_rewrite1(X4) | ~v1_relat_1(X4) | ~v4_rewrite1(X4) | r2_rewrite1(X4,X5,k2_rewrite1(X4,X5)) | ~v1_relat_1(X4) [resolution 128,59] {a:2,w:14,wCS:5600,sine:1,thAx:0,allAx:3,thDist:-3}
154. r2_rewrite1(X4,X5,k2_rewrite1(X4,X5)) | ~v1_relat_1(X4) | ~v4_rewrite1(X4) | ~v2_rewrite1(X4) [duplicate literal removal 153] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
155. r2_rewrite1(X2,k2_rewrite1(X2,X3),X3) | ~v1_relat_1(X2) | ~v4_rewrite1(X2) | ~v2_rewrite1(X2) [duplicate literal removal 152] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
158. ~v1_relat_1(X2) | ~v4_rewrite1(X2) | ~v2_rewrite1(X2) | r2_rewrite1(X2,X3,k2_rewrite1(X2,X4)) | ~r2_rewrite1(X2,X3,X4) | ~v1_relat_1(X2) [resolution 154,64] {a:3,w:18,wCS:7200,sine:1,thAx:0,allAx:4,thDist:-4}
159. r2_rewrite1(X2,X3,k2_rewrite1(X2,X4)) | ~v4_rewrite1(X2) | ~v2_rewrite1(X2) | ~v1_relat_1(X2) | ~r2_rewrite1(X2,X3,X4) [duplicate literal removal 158] {a:3,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
165. ~v1_relat_1(sK0) | ~v4_rewrite1(sK0) | ~v2_rewrite1(sK0) | r2_rewrite1(sK0,k2_rewrite1(sK0,sK1),sK2) <- (2, 5) [resolution 155,132] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
170. ~v4_rewrite1(sK0) | ~v2_rewrite1(sK0) | r2_rewrite1(sK0,k2_rewrite1(sK0,sK1),sK2) <- (2, 5) [subsumption resolution 165,99] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
171. ~v2_rewrite1(sK0) | r2_rewrite1(sK0,k2_rewrite1(sK0,sK1),sK2) <- (2, 3, 5) [subsumption resolution 170,89] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
172. r2_rewrite1(sK0,k2_rewrite1(sK0,sK1),sK2) <- (2, 3, 4, 5) [subsumption resolution 171,94] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
174. 12 <=> r2_rewrite1(sK0,k2_rewrite1(sK0,sK1),sK2) [avatar definition]
176. r2_rewrite1(sK0,k2_rewrite1(sK0,sK1),sK2) <- (12) [avatar component clause 174] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
177. 12 | ~2 | ~3 | ~4 | ~5 [avatar split clause 172,97,92,87,82,174]
183. 13 <=> r3_rewrite1(sK0,k2_rewrite1(sK0,sK1)) [avatar definition]
185. ~r3_rewrite1(sK0,k2_rewrite1(sK0,sK1)) <- (~13) [avatar component clause 183] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
215. ~v4_rewrite1(X0) | ~v2_rewrite1(X0) | ~v1_relat_1(X0) | ~r2_rewrite1(X0,X1,X2) | k2_rewrite1(X0,X2) = X1 | ~r3_rewrite1(X0,k2_rewrite1(X0,X2)) | ~r3_rewrite1(X0,X1) | ~v4_rewrite1(X0) | ~v1_relat_1(X0) [resolution 159,54] {a:4,w:27,wCS:10800,sine:1,thAx:0,allAx:6,thDist:-6}
221. ~v4_rewrite1(X0) | ~v2_rewrite1(X0) | ~v1_relat_1(X0) | ~r2_rewrite1(X0,X1,X2) | k2_rewrite1(X0,X2) = X1 | ~r3_rewrite1(X0,k2_rewrite1(X0,X2)) | ~r3_rewrite1(X0,X1) [duplicate literal removal 215] {a:4,w:23,wCS:9200,sine:1,thAx:0,allAx:6,thDist:-6}
222. k2_rewrite1(X0,X2) = X1 | ~v2_rewrite1(X0) | ~v1_relat_1(X0) | ~r2_rewrite1(X0,X1,X2) | ~v4_rewrite1(X0) | ~r3_rewrite1(X0,X1) [subsumption resolution 221,127] {a:4,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
325. k2_rewrite1(sK0,sK1) != X0 | ~v2_rewrite1(sK0) | ~v1_relat_1(sK0) | ~r2_rewrite1(sK0,X0,sK2) | ~v4_rewrite1(sK0) | ~r3_rewrite1(sK0,X0) <- (~1) [superposition 79,222] {a:5,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
364. k2_rewrite1(sK0,sK1) != X0 | ~v1_relat_1(sK0) | ~r2_rewrite1(sK0,X0,sK2) | ~v4_rewrite1(sK0) | ~r3_rewrite1(sK0,X0) <- (~1, 4) [subsumption resolution 325,94] {a:5,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
365. k2_rewrite1(sK0,sK1) != X0 | ~r2_rewrite1(sK0,X0,sK2) | ~v4_rewrite1(sK0) | ~r3_rewrite1(sK0,X0) <- (~1, 4, 5) [subsumption resolution 364,99] {a:5,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
366. k2_rewrite1(sK0,sK1) != X0 | ~r2_rewrite1(sK0,X0,sK2) | ~r3_rewrite1(sK0,X0) <- (~1, 3, 4, 5) [subsumption resolution 365,89] {a:5,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
405. ~r2_rewrite1(sK0,k2_rewrite1(sK0,sK1),sK2) | ~r3_rewrite1(sK0,k2_rewrite1(sK0,sK1)) <- (~1, 3, 4, 5) [equality resolution 366] {a:6,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
409. ~r3_rewrite1(sK0,k2_rewrite1(sK0,sK1)) <- (~1, 3, 4, 5, 12) [subsumption resolution 405,176] {a:6,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
410. ~13 | 1 | ~3 | ~4 | ~5 | ~12 [avatar split clause 409,174,97,92,87,77,183]
412. ~v2_rewrite1(sK0) | ~v1_relat_1(sK0) | ~v4_rewrite1(sK0) <- (~13) [resolution 185,127] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
418. ~v1_relat_1(sK0) | ~v4_rewrite1(sK0) <- (4, ~13) [subsumption resolution 412,94] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
419. ~v4_rewrite1(sK0) <- (4, 5, ~13) [subsumption resolution 418,99] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
420. $false <- (3, 4, 5, ~13) [subsumption resolution 419,89] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
421. ~3 | ~4 | ~5 | 13 [avatar contradiction clause 420]
425. $false [avatar sat refutation 80,85,90,95,100,177,410,421] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:216,thDist:-216}
% SZS output end Proof for rewrite1__t55_rewrite1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 302
% Active clauses: 50
% Passive clauses: 79
% Final active clauses: 50
% Final passive clauses: 26
% Main loop iterations started: 53
% 
% >>> Simplifying Inferences
% Duplicate literals: 112
% Fw subsumption resolutions: 81
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Equational tautologies: 1
% Forward subsumptions: 55
% 
% >>> Generating Inferences
% Binary resolution: 64
% Unit resulting resolution: 1
% Forward superposition: 24
% Backward superposition: 34
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 6
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 5117
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% term sharing: 0.001 s
% resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.004 s ( own 0.002 s ) 

