i: [1,101,0,2,1,1,0,0] t21_robbins2
a: [3,151,0,2,1,101]
s: 151
i: [1,100,0,2,1,1,0,0] t21_robbins2
a: [3,156,0,2,1,100]
s: 156
i: [1,99,0,2,1,1,0,0] t21_robbins2
a: [3,161,0,2,1,99]
s: 161
i: [1,102,0,4,1,1,0,0] t21_robbins2
a: [3,146,0,4,1,102]
s: 146
i: [1,134,0,2,1,0,0,2] existence_l2_lattices
a: [3,166,0,2,1,134]
s: 166
i: [1,135,0,2,1,0,0,2] existence_l1_robbins1
a: [3,171,0,2,1,135]
s: 171
i: [1,136,0,2,1,0,0,1] existence_l2_robbins1
a: [3,176,0,2,1,136]
s: 176
i: [1,137,0,2,1,0,0,3] existence_l1_struct_0
a: [3,181,0,2,1,137]
s: 181
i: [1,103,0,10,1,1,0,0] t21_robbins2
a: [3,141,0,10,1,103]
s: 141
i: [1,104,0,4,2,0,0,2] dt_l1_robbins1
s: 104
i: [1,105,0,4,2,0,0,2] dt_l2_lattices
s: 105
d: [2,183,1,2,1,1,40,104,171]
a: [3,187,1,2,1,183]
s: 187
i: [1,106,0,4,2,0,0,2] dt_l2_robbins1
s: 106
d: [2,195,1,2,1,1,40,106,151]
a: [3,200,1,2,1,195]
s: 200
i: [1,107,0,4,2,0,0,2] dt_l2_robbins1
s: 107
i: [1,131,0,4,1,0,0,1] existence_m1_subset_1
s: 131
d: [2,213,1,2,1,1,40,107,151]
a: [3,218,1,2,1,213]
s: 218
i: [1,108,0,6,3,0,0,255] cc7_struct_0
s: 108
d: [2,207,2,2,1,1,40,200,105]
a: [3,211,2,2,1,207]
s: 211
i: [1,109,0,6,3,0,0,255] cc5_struct_0
s: 109
i: [1,111,0,6,3,0,0,255] cc2_struct_0
s: 111
d: [2,189,1,2,1,1,40,105,166]
a: [3,193,1,2,1,189]
s: 193
i: [1,112,0,6,3,0,0,255] cc17_struct_0
s: 112
d: [2,196,1,2,1,1,40,106,176]
a: [3,205,1,2,1,196]
s: 205
i: [1,113,0,6,3,0,0,255] cc6_struct_0
s: 113
i: [1,114,0,6,3,0,0,255] cc1_struct_0
s: 114
d: [2,214,1,2,1,1,40,107,176]
a: [3,223,1,2,1,214]
s: 223
i: [1,117,0,6,3,0,0,255] cc4_struct_0
s: 117
d: [2,227,1,4,2,1,40,109,181]
a: [3,233,1,2,1,227]
s: 233
i: [1,115,0,7,3,0,0,255] cc8_struct_0
s: 115
i: [1,118,0,7,3,0,0,255] cc9_struct_0
s: 118
d: [2,228,2,4,2,1,40,109,187]
a: [3,242,2,2,1,228]
s: 242
i: [1,119,0,7,3,0,0,255] cc11_struct_0
s: 119
d: [2,261,2,4,2,1,40,193,111]
a: [3,266,2,2,1,261]
s: 266
i: [1,120,0,7,3,0,0,255] cc11_struct_0
s: 120
i: [1,122,0,7,3,0,0,255] fc9_struct_0
s: 122
d: [2,277,2,2,1,1,40,205,105]
a: [3,281,2,2,1,277]
s: 281
i: [1,123,0,7,3,0,0,255] fc6_struct_0
s: 123
d: [2,287,3,4,2,1,40,281,111]
a: [3,292,3,2,1,287]
s: 292
i: [1,127,0,7,3,0,0,255] fc2_struct_0
s: 127
i: [1,128,0,7,3,0,0,255] fc8_struct_0
s: 128
i: [1,129,0,7,3,0,0,255] fc7_struct_0
s: 129
i: [1,130,0,7,3,0,0,255] fc1_struct_0
s: 130
i: [1,121,0,9,4,0,0,255] cc10_struct_0
s: 121
i: [1,124,0,11,4,0,0,255] dt_o_1_0_robbins2
s: 124
i: [1,132,0,14,4,0,0,1] dt_k3_robbins1
s: 132
i: [1,133,0,19,5,0,0,1] dt_k1_lattices
s: 133
i: [1,125,0,29,6,0,0,1] t19_robbins2
s: 125
i: [1,126,0,31,6,0,0,1] t20_robbins2
s: 126
d: [2,321,1,43,11,0,44,126,125]
d: [2,330,1,33,7,0,18,321]
s: 330
d: [2,313,1,49,11,0,44,125,125]
d: [2,318,1,43,8,0,18,313]
s: 318
d: [2,333,2,30,7,1,44,141,330]
d: [2,344,2,23,6,1,19,333]
d: [2,345,2,19,5,2,21,344,146]
d: [2,346,2,17,4,3,21,345,151]
d: [2,347,2,15,3,4,21,346,156]
d: [2,348,2,13,2,5,21,347,161]
s: 348
d: [2,322,1,55,11,0,44,126,126]
d: [2,329,1,45,7,0,18,322]
s: 329
d: [2,367,3,15,4,5,40,348,132]
d: [2,370,3,13,3,6,21,367,218]
d: [2,371,3,11,2,6,21,370,161]
s: 371
d: [2,323,1,61,11,0,44,125,126]
d: [2,328,1,51,7,0,18,323]
s: 328
d: [2,319,1,61,11,0,44,126,125]
d: [2,332,1,55,8,0,18,319]
s: 332
d: [2,394,4,16,5,6,40,371,133]
d: [2,395,4,12,4,6,18,394]
d: [2,396,4,8,3,6,21,395,146]
d: [2,397,4,6,2,7,21,396,200]
d: [2,398,4,4,1,7,21,397,161]
s: 398
r: 348
r: 371
d: [2,506,5,0,0,7,40,398,131]
e: 506
d: [2,505,5,0,0,7,40,398,146]
e: 505
d: [2,497,5,0,0,7,52,146,398]
e: 497
d: [2,498,5,0,0,7,52,161,156,151,398,124]
e: 498
d: [2,499,5,0,0,7,52,161,218,146,398,132]
e: 499
d: [2,500,5,0,0,7,52,161,200,146,146,398,133]
e: 500
d: [2,501,5,0,0,7,52,161,200,146,131,398,133]
e: 501
d: [2,504,5,6,3,7,40,398,124]
d: [2,514,5,4,2,7,21,504,151]
d: [2,515,5,2,1,7,21,514,156]
d: [2,516,5,0,0,7,21,515,161]
e: 516
f: 516,213,195,99,100,101,102,103
% Refutation found. Thanks to Tanya!
% SZS status Theorem for robbins2__t21_robbins2
% SZS output start Proof for robbins2__t21_robbins2
1. ! [X0] : ((l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k3_robbins1(X0,k1_lattices(X0,X1,X1)) = k3_robbins1(X0,X1))) [input t21_robbins2]
2. ~! [X0] : ((l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k3_robbins1(X0,k1_lattices(X0,X1,X1)) = k3_robbins1(X0,X1))) [negated conjecture 1]
14. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l2_lattices(X0) & ~v2_struct_0(X0)) => m1_subset_1(k1_lattices(X0,X1,X2),u1_struct_0(X0))) [input dt_k1_lattices]
16. ! [X0,X1] : ((m1_subset_1(X1,u1_struct_0(X0)) & l1_robbins1(X0) & ~v2_struct_0(X0)) => m1_subset_1(k3_robbins1(X0,X1),u1_struct_0(X0))) [input dt_k3_robbins1]
20. ! [X0] : (l2_robbins1(X0) => (l1_robbins1(X0) & l2_lattices(X0))) [input dt_l2_robbins1]
22. ! [X0] : ((l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => m1_subset_1(o_1_0_robbins2(X0),u1_struct_0(X0))) [input dt_o_1_0_robbins2]
35. ! [X0] : ((l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,k1_lattices(X0,X1,X2)),k3_robbins1(X0,X2))) = X2))) [input t19_robbins2]
36. ! [X0] : ((l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => k3_robbins1(X0,X1) = k3_robbins1(X0,k1_lattices(X0,X1,k3_robbins1(X0,k1_lattices(X0,X2,k3_robbins1(X0,X1)))))))) [input t20_robbins2]
37. ? [X0] : (? [X1] : (k3_robbins1(X0,k1_lattices(X0,X1,X1)) != k3_robbins1(X0,X1) & m1_subset_1(X1,u1_struct_0(X0))) & (l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
38. ? [X0] : (? [X1] : (k3_robbins1(X0,k1_lattices(X0,X1,X1)) != k3_robbins1(X0,X1) & m1_subset_1(X1,u1_struct_0(X0))) & l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) [flattening 37]
41. ! [X0] : ((l1_robbins1(X0) & l2_lattices(X0)) | ~l2_robbins1(X0)) [ennf transformation 20]
68. ! [X0] : (m1_subset_1(o_1_0_robbins2(X0),u1_struct_0(X0)) | (~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0))) [ennf transformation 22]
69. ! [X0] : (m1_subset_1(o_1_0_robbins2(X0),u1_struct_0(X0)) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0)) [flattening 68]
70. ! [X0] : (! [X1] : (! [X2] : (k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,k1_lattices(X0,X1,X2)),k3_robbins1(X0,X2))) = X2 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0))) [ennf transformation 35]
71. ! [X0] : (! [X1] : (! [X2] : (k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,k1_lattices(X0,X1,X2)),k3_robbins1(X0,X2))) = X2 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0)) [flattening 70]
72. ! [X0] : (! [X1] : (! [X2] : (k3_robbins1(X0,X1) = k3_robbins1(X0,k1_lattices(X0,X1,k3_robbins1(X0,k1_lattices(X0,X2,k3_robbins1(X0,X1))))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0))) [ennf transformation 36]
73. ! [X0] : (! [X1] : (! [X2] : (k3_robbins1(X0,X1) = k3_robbins1(X0,k1_lattices(X0,X1,k3_robbins1(X0,k1_lattices(X0,X2,k3_robbins1(X0,X1))))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0)) [flattening 72]
82. ! [X0,X1] : (m1_subset_1(k3_robbins1(X0,X1),u1_struct_0(X0)) | (~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_robbins1(X0) | v2_struct_0(X0))) [ennf transformation 16]
83. ! [X0,X1] : (m1_subset_1(k3_robbins1(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_robbins1(X0) | v2_struct_0(X0)) [flattening 82]
84. ! [X0,X1,X2] : (m1_subset_1(k1_lattices(X0,X1,X2),u1_struct_0(X0)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | v2_struct_0(X0))) [ennf transformation 14]
85. ! [X0,X1,X2] : (m1_subset_1(k1_lattices(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | v2_struct_0(X0)) [flattening 84]
86. ? [X0] : (? [X1] : (k3_robbins1(X0,k1_lattices(X0,X1,X1)) != k3_robbins1(X0,X1) & m1_subset_1(X1,u1_struct_0(X0))) & l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => (? [X1] : (k3_robbins1(sK0,k1_lattices(sK0,X1,X1)) != k3_robbins1(sK0,X1) & m1_subset_1(X1,u1_struct_0(sK0))) & l2_robbins1(sK0) & v1_robbins2(sK0) & ~v2_struct_0(sK0)) [choice axiom]
87. ? [X1] : (k3_robbins1(sK0,k1_lattices(sK0,X1,X1)) != k3_robbins1(sK0,X1) & m1_subset_1(X1,u1_struct_0(sK0))) => (k3_robbins1(sK0,k1_lattices(sK0,sK1,sK1)) != k3_robbins1(sK0,sK1) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
88. (k3_robbins1(sK0,k1_lattices(sK0,sK1,sK1)) != k3_robbins1(sK0,sK1) & m1_subset_1(sK1,u1_struct_0(sK0))) & l2_robbins1(sK0) & v1_robbins2(sK0) & ~v2_struct_0(sK0) [skolemisation 38,87,86]
99. ~v2_struct_0(sK0) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. v1_robbins2(sK0) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
101. l2_robbins1(sK0) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
102. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 88] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
103. k3_robbins1(sK0,k1_lattices(sK0,sK1,sK1)) != k3_robbins1(sK0,sK1) [cnf transformation 88] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
106. ~l2_robbins1(X0) | l2_lattices(X0) [cnf transformation 41] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
107. ~l2_robbins1(X0) | l1_robbins1(X0) [cnf transformation 41] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
124. m1_subset_1(o_1_0_robbins2(X0),u1_struct_0(X0)) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0) [cnf transformation 69] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
125. k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,k1_lattices(X0,X1,X2)),k3_robbins1(X0,X2))) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0) [cnf transformation 71] {a:0,w:29,wCS:11600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
126. k3_robbins1(X0,X1) = k3_robbins1(X0,k1_lattices(X0,X1,k3_robbins1(X0,k1_lattices(X0,X2,k3_robbins1(X0,X1))))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0) [cnf transformation 73] {a:0,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
132. m1_subset_1(k3_robbins1(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_robbins1(X0) | v2_struct_0(X0) [cnf transformation 83] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
133. m1_subset_1(k1_lattices(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | v2_struct_0(X0) [cnf transformation 85] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
139. 1 <=> k3_robbins1(sK0,k1_lattices(sK0,sK1,sK1)) = k3_robbins1(sK0,sK1) [avatar definition]
141. k3_robbins1(sK0,k1_lattices(sK0,sK1,sK1)) != k3_robbins1(sK0,sK1) <- (~1) [avatar component clause 139] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. ~1 [avatar split clause 103,139]
144. 2 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
146. m1_subset_1(sK1,u1_struct_0(sK0)) <- (2) [avatar component clause 144] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
147. 2 [avatar split clause 102,144]
149. 3 <=> l2_robbins1(sK0) [avatar definition]
151. l2_robbins1(sK0) <- (3) [avatar component clause 149] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
152. 3 [avatar split clause 101,149]
154. 4 <=> v1_robbins2(sK0) [avatar definition]
156. v1_robbins2(sK0) <- (4) [avatar component clause 154] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
157. 4 [avatar split clause 100,154]
159. 5 <=> v2_struct_0(sK0) [avatar definition]
161. ~v2_struct_0(sK0) <- (~5) [avatar component clause 159] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
162. ~5 [avatar split clause 99,159]
195. l2_lattices(sK0) <- (3) [resolution 106,151] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
198. 12 <=> l2_lattices(sK0) [avatar definition]
200. l2_lattices(sK0) <- (12) [avatar component clause 198] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
201. 12 | ~3 [avatar split clause 195,149,198]
213. l1_robbins1(sK0) <- (3) [resolution 107,151] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
216. 15 <=> l1_robbins1(sK0) [avatar definition]
218. l1_robbins1(sK0) <- (15) [avatar component clause 216] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
219. 15 | ~3 [avatar split clause 213,149,216]
321. k3_robbins1(X0,X2) = k3_robbins1(X0,k1_lattices(X0,X2,X2)) | ~m1_subset_1(k3_robbins1(X0,k1_lattices(X0,X1,X2)),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0) [superposition 126,125] {a:1,w:43,wCS:17200,sine:1,thAx:0,allAx:2,thDist:-2}
330. k3_robbins1(X0,X2) = k3_robbins1(X0,k1_lattices(X0,X2,X2)) | ~m1_subset_1(k3_robbins1(X0,k1_lattices(X0,X1,X2)),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) [duplicate literal removal 321] {a:1,w:33,wCS:13200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
333. k3_robbins1(sK0,sK1) != k3_robbins1(sK0,sK1) | ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,X0,sK1)),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l2_robbins1(sK0) | ~v1_robbins2(sK0) | v2_struct_0(sK0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1) [superposition 141,330] {a:2,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
344. ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,X0,sK1)),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l2_robbins1(sK0) | ~v1_robbins2(sK0) | v2_struct_0(sK0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1) [trivial inequality removal 333] {a:2,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
345. ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,X0,sK1)),u1_struct_0(sK0)) | ~l2_robbins1(sK0) | ~v1_robbins2(sK0) | v2_struct_0(sK0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 2) [subsumption resolution 344,146] {a:2,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
346. ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,X0,sK1)),u1_struct_0(sK0)) | ~v1_robbins2(sK0) | v2_struct_0(sK0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 2, 3) [subsumption resolution 345,151] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
347. ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,X0,sK1)),u1_struct_0(sK0)) | v2_struct_0(sK0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 2, 3, 4) [subsumption resolution 346,156] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
348. ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,X0,sK1)),u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 2, 3, 4, ~5) [subsumption resolution 347,161] {a:2,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
367. ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k1_lattices(sK0,X0,sK1),u1_struct_0(sK0)) | ~l1_robbins1(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4, ~5) [resolution 348,132] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
370. ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(k1_lattices(sK0,X0,sK1),u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~1, 2, 3, 4, ~5, 15) [subsumption resolution 367,218] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
371. ~m1_subset_1(k1_lattices(sK0,X0,sK1),u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 2, 3, 4, ~5, 15) [subsumption resolution 370,161] {a:3,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
394. ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l2_lattices(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4, ~5, 15) [resolution 371,133] {a:4,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
395. ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l2_lattices(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4, ~5, 15) [duplicate literal removal 394] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
396. ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l2_lattices(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4, ~5, 15) [subsumption resolution 395,146] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
397. ~m1_subset_1(X0,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~1, 2, 3, 4, ~5, 12, 15) [subsumption resolution 396,200] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
398. ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 2, 3, 4, ~5, 12, 15) [subsumption resolution 397,161] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
504. ~l2_robbins1(sK0) | ~v1_robbins2(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4, ~5, 12, 15) [resolution 398,124] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
514. ~v1_robbins2(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4, ~5, 12, 15) [subsumption resolution 504,151] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
515. v2_struct_0(sK0) <- (~1, 2, 3, 4, ~5, 12, 15) [subsumption resolution 514,156] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
516. $false <- (~1, 2, 3, 4, ~5, 12, 15) [subsumption resolution 515,161] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
517. 1 | ~2 | ~3 | ~4 | 5 | ~12 | ~15 [avatar contradiction clause 516]
518. $false [avatar sat refutation 142,147,152,157,162,201,219,517] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:259,thDist:-259}
% SZS output end Proof for robbins2__t21_robbins2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 35
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 39
% Generated clauses: 289
% Active clauses: 53
% Passive clauses: 113
% Final active clauses: 51
% Final passive clauses: 56
% Main loop iterations started: 57
% 
% >>> Simplifying Inferences
% Duplicate literals: 290
% Trivial inequalities: 1
% Fw subsumption resolutions: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 12
% Forward subsumptions: 28
% Backward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 33
% Unit resulting resolution: 5
% Forward superposition: 38
% Backward superposition: 37
% Self superposition: 11
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 16
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 37
% SAT solver unit clauses: 9
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 9722
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% term sharing: 0.001 s
% non unit clause index maintenance: 0.001 s ( own 0.0000 s ) 
% forward subsumption resolution: 0.001 s
% other: 0.005 s ( own 0.001 s ) 

