i: [1,236,0,2,1,1,0,0] t129_xreal_1
a: [3,375,0,2,1,236]
s: 375
i: [1,235,0,2,1,1,0,0] t129_xreal_1
a: [3,380,0,2,1,235]
s: 380
i: [1,238,0,3,1,1,0,0] t129_xreal_1
a: [3,365,0,3,1,238]
s: 365
i: [1,237,0,3,1,1,0,0] t129_xreal_1
a: [3,370,0,3,1,237]
s: 370
i: [1,239,0,5,1,1,0,0] t129_xreal_1
a: [3,360,0,5,1,239]
s: 360
i: [1,240,0,2,1,0,0,2] fc1_numbers
a: [3,385,0,2,1,240]
s: 385
i: [1,241,0,2,1,0,0,3] spc1_boole
a: [3,390,0,2,1,241]
s: 390
i: [1,242,0,2,1,0,0,2] fc1_xboole_0
a: [3,395,0,2,1,242]
s: 395
i: [1,251,0,2,1,0,0,3] fc6_ordinal1
a: [3,435,0,2,1,251]
s: 435
i: [1,250,0,2,1,0,0,3] fc6_ordinal1
a: [3,440,0,2,1,250]
s: 440
i: [1,252,0,2,1,0,0,2] spc1_numerals
a: [3,460,0,2,1,252]
s: 460
i: [1,324,0,2,1,0,0,2] rc2_xboole_0
a: [3,472,0,2,1,324]
s: 472
i: [1,325,0,2,1,0,0,4] rc4_ordinal1
a: [3,477,0,2,1,325]
s: 477
i: [1,326,0,2,1,0,0,4] rc1_ordinal1
a: [3,482,0,2,1,326]
s: 482
i: [1,327,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,487,0,2,1,327]
s: 487
i: [1,328,0,2,1,0,0,1] rc1_xreal_0
a: [3,492,0,2,1,328]
s: 492
i: [1,329,0,2,1,0,0,2] rc1_xxreal_0
a: [3,497,0,2,1,329]
s: 497
i: [1,330,0,2,1,0,0,2] rc1_xboole_0
a: [3,502,0,2,1,330]
s: 502
i: [1,334,0,2,1,0,0,255] rc2_ordinal1
a: [3,507,0,2,1,334]
s: 507
i: [1,333,0,2,1,0,0,255] rc2_ordinal1
a: [3,512,0,2,1,333]
s: 512
i: [1,332,0,2,1,0,0,255] rc2_ordinal1
a: [3,517,0,2,1,332]
s: 517
i: [1,331,0,2,1,0,0,255] rc2_ordinal1
a: [3,522,0,2,1,331]
s: 522
i: [1,336,0,2,1,0,0,4] rc5_ordinal1
a: [3,527,0,2,1,336]
s: 527
i: [1,335,0,2,1,0,0,4] rc5_ordinal1
a: [3,532,0,2,1,335]
s: 532
i: [1,338,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,537,0,2,1,338]
s: 537
i: [1,337,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,542,0,2,1,337]
s: 542
i: [1,342,0,2,1,0,0,2] rc3_xreal_0
a: [3,547,0,2,1,342]
s: 547
i: [1,341,0,2,1,0,0,2] rc3_xreal_0
a: [3,552,0,2,1,341]
s: 552
i: [1,340,0,2,1,0,0,2] rc3_xreal_0
a: [3,557,0,2,1,340]
s: 557
i: [1,339,0,2,1,0,0,2] rc3_xreal_0
a: [3,562,0,2,1,339]
s: 562
i: [1,346,0,2,1,0,0,2] rc2_xreal_0
a: [3,567,0,2,1,346]
s: 567
i: [1,345,0,2,1,0,0,2] rc2_xreal_0
a: [3,572,0,2,1,345]
s: 572
i: [1,344,0,2,1,0,0,2] rc2_xreal_0
a: [3,577,0,2,1,344]
s: 577
i: [1,343,0,2,1,0,0,2] rc2_xreal_0
a: [3,582,0,2,1,343]
s: 582
i: [1,348,0,2,1,0,0,2] rc3_xxreal_0
a: [3,587,0,2,1,348]
s: 587
i: [1,347,0,2,1,0,0,2] rc3_xxreal_0
a: [3,592,0,2,1,347]
s: 592
i: [1,350,0,2,1,0,0,2] rc2_xxreal_0
a: [3,597,0,2,1,350]
s: 597
i: [1,349,0,2,1,0,0,2] rc2_xxreal_0
a: [3,602,0,2,1,349]
s: 602
i: [1,354,0,2,1,0,0,2] rc4_xreal_0
a: [3,607,0,2,1,354]
s: 607
i: [1,353,0,2,1,0,0,2] rc4_xreal_0
a: [3,612,0,2,1,353]
s: 612
i: [1,352,0,2,1,0,0,2] rc4_xreal_0
a: [3,617,0,2,1,352]
s: 617
i: [1,351,0,2,1,0,0,2] rc4_xreal_0
a: [3,622,0,2,1,351]
s: 622
i: [1,356,0,2,1,0,0,2] rc4_xxreal_0
a: [3,627,0,2,1,356]
s: 627
i: [1,355,0,2,1,0,0,2] rc4_xxreal_0
a: [3,632,0,2,1,355]
s: 632
i: [1,307,0,3,1,0,0,255] reflexivity_r1_tarski
s: 307
i: [1,243,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,400,0,3,1,243]
s: 400
i: [1,244,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,405,0,3,1,244]
s: 405
r: 395
d: [2,634,0,2,1,2,23,395,405]
a: [3,638,0,2,1,634]
s: 638
i: [1,245,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,410,0,3,1,245]
s: 410
i: [1,246,0,3,1,0,0,2] t1_numerals
a: [3,415,0,3,1,246]
s: 415
i: [1,255,0,3,1,0,0,2] spc1_numerals
a: [3,445,0,3,1,255]
s: 445
i: [1,254,0,3,1,0,0,2] spc1_numerals
a: [3,450,0,3,1,254]
s: 450
d: [2,640,0,3,1,2,22,415,405]
a: [3,644,0,3,1,640]
s: 644
d: [2,646,0,3,1,2,22,450,400]
a: [3,650,0,3,1,646]
s: 650
i: [1,257,0,4,2,0,0,4] cc6_ordinal1
s: 257
i: [1,258,0,4,2,0,0,4] cc2_xcmplx_0
s: 258
d: [2,652,1,2,1,1,40,257,477]
a: [3,657,1,2,1,652]
s: 657
i: [1,259,0,4,2,0,0,4] cc2_xreal_0
s: 259
d: [2,653,1,2,1,1,40,257,527]
a: [3,662,1,2,1,653]
s: 662
i: [1,260,0,4,2,0,0,4] cc2_xxreal_0
s: 260
i: [1,261,0,4,2,0,0,255] cc1_ordinal1
s: 261
d: [2,664,1,2,1,1,40,258,477]
a: [3,669,1,2,1,664]
s: 669
i: [1,262,0,4,2,0,0,255] cc1_ordinal1
s: 262
d: [2,665,1,2,1,1,40,258,527]
a: [3,674,1,2,1,665]
s: 674
i: [1,266,0,4,2,0,0,2] cc3_xreal_0
s: 266
i: [1,267,0,4,2,0,0,2] cc4_xreal_0
s: 267
d: [2,750,1,2,1,1,40,266,380]
a: [3,759,1,2,1,750]
s: 759
i: [1,276,0,4,2,0,0,255] cc9_ordinal1
s: 276
d: [2,751,1,2,1,1,40,266,375]
a: [3,764,1,2,1,751]
s: 764
i: [1,277,0,4,2,0,0,4] cc7_ordinal1
s: 277
i: [1,278,0,4,2,0,0,4] cc3_ordinal1
s: 278
d: [2,771,1,2,1,1,40,267,380]
a: [3,780,1,2,1,771]
s: 780
i: [1,306,0,4,1,0,0,3] existence_m1_subset_1
s: 306
d: [2,772,1,2,1,1,40,267,375]
a: [3,785,1,2,1,772]
s: 785
i: [1,247,0,4,1,0,0,2] dt_k5_numbers
a: [3,420,0,4,1,247]
s: 420
i: [1,248,0,4,1,0,0,1] dt_k6_numbers
a: [3,425,0,4,1,248]
s: 425
d: [2,676,1,2,1,1,40,259,477]
a: [3,681,1,2,1,676]
s: 681
i: [1,253,0,4,1,0,0,2] spc1_numerals
a: [3,455,0,4,1,253]
s: 455
d: [2,677,1,2,1,1,40,259,527]
a: [3,686,1,2,1,677]
s: 686
d: [2,864,0,4,1,2,22,420,400]
a: [3,868,0,4,1,864]
s: 868
d: [2,870,0,4,1,2,22,425,400]
a: [3,874,0,4,1,870]
s: 874
d: [2,688,1,2,1,1,40,260,477]
a: [3,693,1,2,1,688]
s: 693
d: [2,879,0,4,1,2,22,455,400]
a: [3,883,0,4,1,879]
s: 883
d: [2,689,1,2,1,1,40,260,527]
a: [3,698,1,2,1,689]
s: 698
i: [1,279,0,5,2,0,0,2] t6_boole
s: 279
i: [1,280,0,5,2,0,0,3] cc8_ordinal1
s: 280
d: [2,700,1,2,1,1,40,261,435]
a: [3,708,1,2,1,700]
s: 708
i: [1,281,0,5,2,0,0,2] cc1_xcmplx_0
s: 281
d: [2,702,1,2,1,1,40,261,482]
a: [3,718,1,2,1,702]
s: 718
i: [1,282,0,5,2,0,0,2] cc1_xreal_0
s: 282
i: [1,303,0,5,2,0,0,255] t7_boole
s: 303
d: [2,725,1,2,1,1,40,262,435]
a: [3,733,1,2,1,725]
s: 733
i: [1,249,0,5,1,0,0,3] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,430,0,5,1,249]
s: 430
d: [2,727,1,2,1,1,40,262,482]
a: [3,743,1,2,1,727]
s: 743
i: [1,321,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,467,0,5,2,321]
s: 467
d: [2,888,0,5,2,1,22,279,405]
s: 888
d: [2,929,1,5,2,2,44,365,888]
s: 929
i: [1,287,0,6,3,0,0,255] cc2_ordinal1
s: 287
d: [2,934,2,4,2,3,40,929,467]
d: [2,935,2,2,1,4,21,934,785]
a: [3,939,2,2,1,935]
s: 939
i: [1,290,0,6,3,0,0,2] cc5_xxreal_0
s: 290
i: [1,292,0,6,3,0,0,2] cc5_xxreal_0
s: 292
d: [2,930,1,5,2,2,44,370,888]
s: 930
i: [1,293,0,6,3,0,0,2] cc3_xxreal_0
s: 293
d: [2,982,2,4,2,3,40,930,467]
d: [2,983,2,2,1,4,21,982,780]
a: [3,987,2,2,1,983]
s: 987
i: [1,295,0,6,3,0,0,2] cc3_xxreal_0
s: 295
i: [1,297,0,6,3,0,0,2] cc7_xxreal_0
s: 297
d: [2,928,1,7,2,2,44,360,888]
s: 928
i: [1,298,0,6,3,0,0,2] cc7_xxreal_0
s: 298
d: [2,752,1,2,1,1,40,266,492]
a: [3,769,1,2,1,752]
s: 769
i: [1,299,0,6,2,0,0,255] t1_subset
s: 299
i: [1,308,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 308
d: [2,773,1,2,1,1,40,267,492]
a: [3,790,1,2,1,773]
s: 790
i: [1,263,0,7,3,0,0,4] cc5_ordinal1
s: 263
d: [2,792,1,2,1,1,40,276,638]
a: [3,799,1,2,1,792]
s: 799
i: [1,264,0,7,2,0,0,1] t2_arithm
s: 264
i: [1,265,0,7,2,0,0,3] t3_arithm
s: 265
d: [2,793,1,2,1,1,40,276,502]
a: [3,804,1,2,1,793]
s: 804
i: [1,301,0,7,2,0,0,255] t3_subset
s: 301
d: [2,794,1,2,1,1,40,276,622]
a: [3,809,1,2,1,794]
s: 809
i: [1,302,0,7,3,0,0,2] t8_boole
s: 302
i: [1,256,0,8,3,0,0,255] cc10_ordinal1
s: 256
d: [2,795,1,2,1,1,40,276,632]
a: [3,814,1,2,1,795]
s: 814
i: [1,284,0,8,4,0,0,2] cc4_xxreal_0
s: 284
r: 467
r: 939
r: 987
a: [3,938,2,2,1,935]
s: 938
r: 938
a: [3,464,0,2,1,321]
s: 464
r: 290
r: 292
r: 293
r: 284
d: [2,1152,1,0,0,2,40,464,627]
e: 1152
d: [2,1151,1,0,0,2,40,464,612]
e: 1151
d: [2,1150,1,0,0,2,40,464,602]
e: 1150
d: [2,1149,1,0,0,2,40,464,592]
e: 1149
d: [2,1148,1,0,0,2,40,464,577]
e: 1148
d: [2,1147,1,0,0,2,40,464,557]
e: 1147
d: [2,1146,2,0,0,2,40,464,698]
e: 1146
d: [2,1145,1,0,0,2,40,464,497]
e: 1145
d: [2,1144,2,0,0,2,40,464,790]
e: 1144
d: [2,1143,2,0,0,2,40,464,693]
e: 1143
d: [2,1142,2,0,0,2,40,464,785]
e: 1142
d: [2,1141,2,0,0,2,40,464,780]
e: 1141
d: [2,1138,2,0,0,2,52,780,464]
e: 1138
d: [2,1139,1,0,0,2,52,380,464,267]
e: 1139
d: [2,1140,1,0,0,2,52,477,464,260]
e: 1140
r: 464
s: 467
s: 939
s: 290
s: 987
s: 292
s: 293
d: [2,1037,2,6,3,1,40,284,780]
d: [2,1049,2,4,2,2,21,1037,987]
a: [3,1053,2,2,1,1049]
s: 1053
s: 284
a: [3,1056,2,2,1,1049]
s: 1056
i: [1,286,0,8,4,0,0,2] cc6_xxreal_0
s: 286
i: [1,288,0,8,4,0,0,2] cc8_xxreal_0
s: 288
d: [2,1038,2,6,3,1,40,284,785]
d: [2,1059,2,4,2,2,21,1038,939]
a: [3,1063,2,2,1,1059]
s: 1063
i: [1,300,0,8,3,0,0,255] t2_subset
s: 300
a: [3,1066,2,2,1,1059]
s: 1066
i: [1,318,0,8,3,0,0,1] fc3_xcmplx_0
s: 318
i: [1,320,0,8,3,0,0,1] fc6_xreal_0
s: 320
d: [2,816,1,2,1,1,40,277,638]
a: [3,823,1,2,1,816]
s: 823
i: [1,305,0,9,3,0,0,255] t5_subset
s: 305
d: [2,817,1,2,1,1,40,277,502]
a: [3,828,1,2,1,817]
s: 828
i: [1,304,0,10,3,0,0,255] t4_subset
s: 304
i: [1,322,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 322
d: [2,1314,1,7,3,1,40,322,365]
d: [2,1320,1,5,2,2,21,1314,785]
a: [3,1324,1,3,1,1320]
s: 1324
i: [1,270,0,11,5,0,0,1] t2_real
s: 270
d: [2,1315,1,7,3,1,40,322,370]
d: [2,1326,1,5,2,2,21,1315,780]
a: [3,1330,1,3,1,1326]
s: 1330
i: [1,271,0,11,5,0,0,1] t1_real
s: 271
i: [1,272,0,11,5,0,0,1] t8_real
s: 272
d: [2,1338,1,12,4,1,40,270,360]
a: [3,1348,1,4,1,1338]
s: 1348
d: [2,1396,2,0,0,3,52,380,375,1348,320]
e: 1396
d: [2,1397,2,4,2,1,40,1348,320]
d: [2,1399,2,2,1,2,21,1397,375]
d: [2,1400,2,0,0,3,21,1399,380]
e: 1400
r: 1348
i: [1,273,0,11,5,0,0,1] t7_real
s: 273
a: [3,1347,1,4,1,1338]
s: 1347
i: [1,274,0,11,5,0,0,1] t3_real
s: 274
i: [1,275,0,11,5,0,0,1] t4_real
s: 275
a: [3,1352,1,4,1,1338]
s: 1352
i: [1,319,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 319
d: [2,1368,1,12,4,1,40,271,360]
a: [3,1382,1,4,1,1368]
s: 1382
i: [1,309,0,12,5,0,0,1] fc26_xreal_0
s: 309
d: [2,1456,3,0,0,5,52,1056,380,1066,375,1352,309]
e: 1456
d: [2,1457,2,8,4,1,40,309,1352]
d: [2,1465,2,6,3,2,21,1457,375]
d: [2,1466,2,4,2,3,21,1465,1066]
d: [2,1467,2,2,1,4,21,1466,380]
d: [2,1468,2,0,0,5,21,1467,1056]
e: 1468
r: 1053
r: 1056
r: 1063
r: 1066
r: 1352
i: [1,310,0,12,5,0,0,1] fc24_xreal_0
s: 310
a: [3,1356,1,2,1,1338]
s: 1356
i: [1,311,0,12,5,0,0,1] fc23_xreal_0
s: 311
a: [3,1377,1,2,1,1368]
s: 1377
i: [1,312,0,12,5,0,0,1] fc25_xreal_0
s: 312
i: [1,317,0,12,5,0,0,1] fc8_xcmplx_0
s: 317
s: 1053
i: [1,268,0,13,6,0,0,1] t5_real
s: 268
s: 1056
i: [1,269,0,13,6,0,0,1] t6_real
s: 269
i: [1,316,0,14,4,0,0,2] existence_m2_subset_1
s: 316
s: 1063
i: [1,313,0,15,5,0,0,2] dt_m2_subset_1
s: 313
s: 1066
i: [1,314,0,15,5,0,0,2] redefinition_m2_subset_1
s: 314
i: [1,315,0,15,5,0,0,2] redefinition_m2_subset_1
s: 315
a: [3,1351,1,4,1,1338]
s: 1351
i: [1,323,0,17,4,0,0,2] spc7_arithm
s: 323
d: [2,1402,2,4,1,1,40,1347,267]
a: [3,1407,2,4,1,1402]
s: 1407
d: [2,818,1,2,1,1,40,277,622]
a: [3,833,1,2,1,818]
s: 833
d: [2,819,1,2,1,1,40,277,632]
a: [3,838,1,2,1,819]
s: 838
d: [2,1403,2,4,1,1,40,1347,266]
a: [3,1412,2,4,1,1403]
s: 1412
d: [2,840,1,2,1,1,40,278,638]
a: [3,847,1,2,1,840]
s: 847
d: [2,1414,2,8,2,1,40,1352,290]
a: [3,1417,2,4,1,1414]
s: 1417
d: [2,1701,3,0,0,5,52,987,759,939,764,1417,317]
e: 1701
d: [2,1702,3,8,4,1,40,1417,317]
d: [2,1707,3,6,3,2,21,1702,764]
d: [2,1708,3,4,2,3,21,1707,939]
d: [2,1709,3,2,1,4,21,1708,759]
d: [2,1710,3,0,0,5,21,1709,987]
e: 1710
r: 1063
r: 1066
r: 1351
r: 1356
r: 1417
a: [3,1355,1,2,1,1338]
s: 1355
r: 1355
r: 1377
r: 1382
s: 1356
a: [3,1378,1,2,1,1368]
s: 1378
d: [2,889,1,2,1,1,40,280,644]
d: [2,1168,2,4,2,3,40,467,929]
d: [2,1170,2,2,1,4,21,1168,785]
d: [2,1169,2,4,2,3,40,467,930]
d: [2,1172,2,2,1,4,21,1169,780]
d: [2,1267,2,2,1,1,40,823,260]
d: [2,1268,2,2,1,1,40,823,259]
d: [2,1389,1,8,4,1,40,272,370]
d: [2,1475,1,6,3,2,21,1389,380]
d: [2,1388,1,8,4,1,40,272,365]
d: [2,1486,1,6,3,2,21,1388,375]
d: [2,1490,2,4,2,3,21,1465,380]
d: [2,1535,3,4,2,1,40,1053,292]
d: [2,1536,3,2,1,2,21,1535,780]
d: [2,1578,3,4,2,1,40,1063,292]
d: [2,1579,3,2,1,2,21,1578,785]
d: [2,1720,2,4,2,1,40,1355,290]
d: [2,1722,2,2,1,2,21,1720,638]
d: [2,1565,1,16,5,1,40,269,360]
d: [2,1728,1,12,4,2,21,1565,1347]
d: [2,1732,2,4,2,1,40,1378,293]
d: [2,1735,2,2,1,2,21,1732,638]
f: 1735,1728,1722,1710,1579,1536,1490,1486,1475,1400,1268,1267,1172,1170,1152,889,772,771,751,750,640,634,356,321,246,244,242,235,236,237,238,239
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t129_xreal_1
% SZS output start Proof for xreal_1__t129_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers)))) [input t129_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers)))) [negated conjecture 1]
12. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
13. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
16. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
18. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
22. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k3_xcmplx_0(X0,X1))) [input fc26_xreal_0]
29. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
37. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0))))) [input t6_real]
39. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
52. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
59. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
63. ! [X0,X1] : ((v1_xcmplx_0(X1) & ~v1_xboole_0(X1) & v1_xcmplx_0(X0) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k3_xcmplx_0(X0,X1))) [input fc8_xcmplx_0]
72. ? [X0] : (v1_xxreal_0(X0) & v1_xboole_0(X0)) [input rc4_xxreal_0]
79. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
83. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => r1_xxreal_0(X0,X0)) [input reflexivity_r1_xxreal_0]
85. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
88. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
89. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
90. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k3_xcmplx_0(X0,X1))) [input fc6_xreal_0]
107. ? [X0] : (? [X1] : ((r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
108. ? [X0] : (? [X1] : (r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 107]
112. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
113. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 13]
118. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 88]
119. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 89]
122. ! [X0] : (! [X1] : (((v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0)) | ~r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 37]
123. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 122]
128. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 39]
129. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 128]
139. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 79]
140. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 18]
151. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 16]
152. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 151]
153. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 52]
154. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 153]
167. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 22]
168. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 167]
181. ! [X0,X1] : (~v1_xboole_0(k3_xcmplx_0(X0,X1)) | (~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0))) [ennf transformation 63]
182. ! [X0,X1] : (~v1_xboole_0(k3_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0)) [flattening 181]
187. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 90]
188. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 187]
189. ! [X0,X1] : (r1_xxreal_0(X0,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 83]
190. ! [X0,X1] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 189]
195. ? [X0] : (? [X1] : (r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (r1_xxreal_0(k3_xcmplx_0(sK0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
196. ? [X1] : (r1_xxreal_0(k3_xcmplx_0(sK0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) => (r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) & ~r1_xxreal_0(sK1,k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) [choice axiom]
197. (r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) & ~r1_xxreal_0(sK1,k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 108,196,195]
233. ? [X0] : (v1_xxreal_0(X0) & v1_xboole_0(X0)) => (v1_xxreal_0(sK19) & v1_xboole_0(sK19)) [choice axiom]
234. v1_xxreal_0(sK19) & v1_xboole_0(sK19) [skolemisation 72,233]
235. v1_xreal_0(sK0) [cnf transformation 197] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
236. v1_xreal_0(sK1) [cnf transformation 197] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
237. ~r1_xxreal_0(sK0,k6_numbers) [cnf transformation 197] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
238. ~r1_xxreal_0(sK1,k6_numbers) [cnf transformation 197] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
239. r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) [cnf transformation 197] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
242. v1_xboole_0(k1_xboole_0) [cnf transformation 59] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
244. k6_numbers = k1_xboole_0 [cnf transformation 85] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
246. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 29] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
259. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 112] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
260. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 113] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
266. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 118] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
267. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 119] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
269. ~r1_xxreal_0(X0,X1) | v2_xxreal_0(X1) | v1_xboole_0(X0) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 123] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
272. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 129] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
279. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 139] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
280. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 140] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
290. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 152] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
292. ~v2_xxreal_0(X0) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 152] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
293. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 154] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
309. ~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 168] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
317. ~v1_xboole_0(k3_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) [cnf transformation 182] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
320. v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 188] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
321. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 190] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
356. v1_xxreal_0(sK19) [cnf transformation 234] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
358. 1 <=> r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) [avatar definition]
360. r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) <- (1) [avatar component clause 358] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
361. 1 [avatar split clause 239,358]
363. 2 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
365. ~r1_xxreal_0(sK1,k6_numbers) <- (~2) [avatar component clause 363] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
366. ~2 [avatar split clause 238,363]
368. 3 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
370. ~r1_xxreal_0(sK0,k6_numbers) <- (~3) [avatar component clause 368] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
371. ~3 [avatar split clause 237,368]
373. 4 <=> v1_xreal_0(sK1) [avatar definition]
375. v1_xreal_0(sK1) <- (4) [avatar component clause 373] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
376. 4 [avatar split clause 236,373]
378. 5 <=> v1_xreal_0(sK0) [avatar definition]
380. v1_xreal_0(sK0) <- (5) [avatar component clause 378] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
381. 5 [avatar split clause 235,378]
393. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
395. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 393] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
396. 8 [avatar split clause 242,393]
403. 10 <=> k6_numbers = k1_xboole_0 [avatar definition]
405. k6_numbers = k1_xboole_0 <- (10) [avatar component clause 403] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
406. 10 [avatar split clause 244,403]
413. 12 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
415. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (12) [avatar component clause 413] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
416. 12 [avatar split clause 246,413]
463. 22 <=> ! [X1] : ~v1_xxreal_0(X1) [avatar definition]
464. ~v1_xxreal_0(X1) <- (22) [avatar component clause 463] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
466. 23 <=> ! [X0] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0)) [avatar definition]
467. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) <- (23) [avatar component clause 466] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
468. 22 | 23 [avatar split clause 321,466,463]
625. 55 <=> v1_xxreal_0(sK19) [avatar definition]
627. v1_xxreal_0(sK19) <- (55) [avatar component clause 625] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
628. 55 [avatar split clause 356,625]
634. v1_xboole_0(k6_numbers) <- (8, 10) [backward demodulation 395,405] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
636. 57 <=> v1_xboole_0(k6_numbers) [avatar definition]
638. v1_xboole_0(k6_numbers) <- (57) [avatar component clause 636] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
639. 57 | ~8 | ~10 [avatar split clause 634,403,393,636]
640. m1_subset_1(k6_numbers,k4_ordinal1) <- (10, 12) [forward demodulation 415,405] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
642. 58 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
644. m1_subset_1(k6_numbers,k4_ordinal1) <- (58) [avatar component clause 642] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
645. 58 | ~10 | ~12 [avatar split clause 640,413,403,642]
750. v1_xcmplx_0(sK0) <- (5) [resolution 266,380] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
751. v1_xcmplx_0(sK1) <- (4) [resolution 266,375] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
757. 76 <=> v1_xcmplx_0(sK0) [avatar definition]
759. v1_xcmplx_0(sK0) <- (76) [avatar component clause 757] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
760. 76 | ~5 [avatar split clause 750,378,757]
762. 77 <=> v1_xcmplx_0(sK1) [avatar definition]
764. v1_xcmplx_0(sK1) <- (77) [avatar component clause 762] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
765. 77 | ~4 [avatar split clause 751,373,762]
771. v1_xxreal_0(sK0) <- (5) [resolution 267,380] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
772. v1_xxreal_0(sK1) <- (4) [resolution 267,375] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
778. 79 <=> v1_xxreal_0(sK0) [avatar definition]
780. v1_xxreal_0(sK0) <- (79) [avatar component clause 778] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
781. 79 | ~5 [avatar split clause 771,378,778]
783. 80 <=> v1_xxreal_0(sK1) [avatar definition]
785. v1_xxreal_0(sK1) <- (80) [avatar component clause 783] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
786. 80 | ~4 [avatar split clause 772,373,783]
821. 86 <=> v7_ordinal1(k6_numbers) [avatar definition]
823. v7_ordinal1(k6_numbers) <- (86) [avatar component clause 821] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
888. k6_numbers = X0 | ~v1_xboole_0(X0) <- (10) [forward demodulation 279,405] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
889. v7_ordinal1(k6_numbers) <- (58) [resolution 280,644] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
892. 86 | ~58 [avatar split clause 889,642,821]
929. ~r1_xxreal_0(sK1,X0) | ~v1_xboole_0(X0) <- (~2, 10) [superposition 365,888] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
930. ~r1_xxreal_0(sK0,X0) | ~v1_xboole_0(X0) <- (~3, 10) [superposition 370,888] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
937. 103 <=> v1_xboole_0(sK1) [avatar definition]
939. ~v1_xboole_0(sK1) <- (~103) [avatar component clause 937] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
985. 110 <=> v1_xboole_0(sK0) [avatar definition]
987. ~v1_xboole_0(sK0) <- (~110) [avatar component clause 985] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1051. 117 <=> v2_xxreal_0(sK0) [avatar definition]
1053. v2_xxreal_0(sK0) <- (117) [avatar component clause 1051] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1055. 118 <=> v3_xxreal_0(sK0) [avatar definition]
1061. 119 <=> v2_xxreal_0(sK1) [avatar definition]
1063. v2_xxreal_0(sK1) <- (119) [avatar component clause 1061] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1065. 120 <=> v3_xxreal_0(sK1) [avatar definition]
1152. $false <- (22, 55) [resolution 464,627] {a:1,w:0,sine:255,thAx:0,allAx:3,thDist:-3}
1153. ~22 | ~55 [avatar contradiction clause 1152]
1168. ~v1_xxreal_0(sK1) | ~v1_xboole_0(sK1) <- (~2, 10, 23) [resolution 467,929] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1169. ~v1_xxreal_0(sK0) | ~v1_xboole_0(sK0) <- (~3, 10, 23) [resolution 467,930] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1170. ~v1_xboole_0(sK1) <- (~2, 10, 23, 80) [subsumption resolution 1168,785] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1171. ~103 | 2 | ~10 | ~23 | ~80 [avatar split clause 1170,783,466,403,363,937]
1172. ~v1_xboole_0(sK0) <- (~3, 10, 23, 79) [subsumption resolution 1169,780] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1173. ~110 | 3 | ~10 | ~23 | ~79 [avatar split clause 1172,778,466,403,368,985]
1267. v1_xxreal_0(k6_numbers) <- (86) [resolution 823,260] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1268. v1_xreal_0(k6_numbers) <- (86) [resolution 823,259] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1273. 140 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1276. 140 | ~86 [avatar split clause 1267,821,1273]
1278. 141 <=> v1_xreal_0(k6_numbers) [avatar definition]
1281. 141 | ~86 [avatar split clause 1268,821,1278]
1346. 148 <=> v1_xreal_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
1347. v1_xreal_0(k3_xcmplx_0(sK0,sK1)) <- (148) [avatar component clause 1346] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1348. ~v1_xreal_0(k3_xcmplx_0(sK0,sK1)) <- (~148) [avatar component clause 1346] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1350. 149 <=> v3_xxreal_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
1352. v3_xxreal_0(k3_xcmplx_0(sK0,sK1)) <- (149) [avatar component clause 1350] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1354. 150 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1355. v3_xxreal_0(k6_numbers) <- (150) [avatar component clause 1354] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1376. 152 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1378. v2_xxreal_0(k6_numbers) <- (152) [avatar component clause 1376] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1388. v3_xxreal_0(k6_numbers) | v2_xxreal_0(sK1) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (~2) [resolution 272,365] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1389. v3_xxreal_0(k6_numbers) | v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (~3) [resolution 272,370] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1397. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~148) [resolution 1348,320] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1399. ~v1_xreal_0(sK0) <- (4, ~148) [subsumption resolution 1397,375] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1400. $false <- (4, 5, ~148) [subsumption resolution 1399,380] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1401. ~4 | ~5 | 148 [avatar contradiction clause 1400]
1416. 157 <=> v1_xboole_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
1417. v1_xboole_0(k3_xcmplx_0(sK0,sK1)) <- (157) [avatar component clause 1416] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1457. ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (149) [resolution 309,1352] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1465. v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (4, 149) [subsumption resolution 1457,375] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1475. v3_xxreal_0(k6_numbers) | v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (~3, 5) [subsumption resolution 1389,380] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1476. ~141 | 117 | 150 | 3 | ~5 [avatar split clause 1475,378,368,1354,1051,1278]
1486. v3_xxreal_0(k6_numbers) | v2_xxreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (~2, 4) [subsumption resolution 1388,375] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1487. ~141 | 119 | 150 | 2 | ~4 [avatar split clause 1486,373,363,1354,1061,1278]
1490. v3_xxreal_0(sK1) | v3_xxreal_0(sK0) <- (4, 5, 149) [subsumption resolution 1465,380] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1491. 118 | 120 | ~4 | ~5 | ~149 [avatar split clause 1490,1350,378,373,1065,1055]
1535. ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (117) [resolution 1053,292] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1536. ~v3_xxreal_0(sK0) <- (79, 117) [subsumption resolution 1535,780] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1537. ~118 | ~79 | ~117 [avatar split clause 1536,1051,778,1055]
1565. v2_xxreal_0(k6_numbers) | v1_xboole_0(k3_xcmplx_0(sK0,sK1)) | v3_xxreal_0(k3_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k3_xcmplx_0(sK0,sK1)) <- (1) [resolution 269,360] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1578. ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (119) [resolution 1063,292] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1579. ~v3_xxreal_0(sK1) <- (80, 119) [subsumption resolution 1578,785] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1580. ~120 | ~80 | ~119 [avatar split clause 1579,1061,783,1065]
1702. ~v1_xcmplx_0(sK1) | v1_xboole_0(sK1) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (157) [resolution 1417,317] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1707. v1_xboole_0(sK1) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (77, 157) [subsumption resolution 1702,764] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1708. ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (77, ~103, 157) [subsumption resolution 1707,939] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1709. v1_xboole_0(sK0) <- (76, 77, ~103, 157) [subsumption resolution 1708,759] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1710. $false <- (76, 77, ~103, ~110, 157) [subsumption resolution 1709,987] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1711. ~76 | ~77 | 103 | 110 | ~157 [avatar contradiction clause 1710]
1720. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (150) [resolution 1355,290] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1722. ~v1_xxreal_0(k6_numbers) <- (57, 150) [subsumption resolution 1720,638] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1723. ~140 | ~57 | ~150 [avatar split clause 1722,1354,636,1273]
1728. v2_xxreal_0(k6_numbers) | v1_xboole_0(k3_xcmplx_0(sK0,sK1)) | v3_xxreal_0(k3_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (1, 148) [subsumption resolution 1565,1347] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1729. ~141 | 149 | 157 | 152 | ~1 | ~148 [avatar split clause 1728,1346,358,1376,1416,1350,1278]
1732. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (152) [resolution 1378,293] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1735. ~v1_xxreal_0(k6_numbers) <- (57, 152) [subsumption resolution 1732,638] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1736. ~140 | ~57 | ~152 [avatar split clause 1735,1376,636,1273]
1737. $false [avatar sat refutation 361,366,371,376,381,396,406,416,468,628,639,645,760,765,781,786,892,1153,1171,1173,1276,1281,1401,1476,1487,1491,1537,1580,1711,1723,1729,1736] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1301,thDist:-1301}
% SZS output end Proof for xreal_1__t129_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 94
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 122
% Generated clauses: 802
% Active clauses: 173
% Passive clauses: 302
% Final active clauses: 149
% Final passive clauses: 103
% Main loop iterations started: 187
% 
% >>> Simplifying Inferences
% Duplicate literals: 67
% Fw subsumption resolutions: 101
% Fw demodulations: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 30
% Forward subsumptions: 89
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 226
% Unit resulting resolution: 6
% Factoring: 1
% Forward superposition: 54
% Backward superposition: 24
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 62
% Split components: 156
% Sat splitting refutations: 21
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 237
% SAT solver unit clauses: 57
% SAT solver binary clauses: 89
% 
% Memory used [KB]: 5884
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% congruence closure: 0.001 s
% forward subsumption resolution: 0.001 s
% resolution: 0.001 s
% other: 0.008 s ( own 0.004 s ) 

