i: [1,69,0,3,1,1,0,0] t3_subset_1
a: [3,102,0,3,1,69]
s: 102
i: [1,66,0,4,1,1,0,0] t3_subset_1
a: [3,107,0,4,1,66]
s: 107
i: [1,65,0,4,1,1,0,0] t3_subset_1
a: [3,112,0,4,1,65]
s: 112
i: [1,70,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,117,0,2,1,70]
s: 117
i: [1,71,0,2,1,0,0,255] fc1_xboole_0
a: [3,122,0,2,1,71]
s: 122
i: [1,92,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,132,0,2,1,92]
s: 132
i: [1,93,0,2,1,0,0,2] rc2_xboole_0
a: [3,137,0,2,1,93]
s: 137
i: [1,94,0,2,1,0,0,2] rc1_xboole_0
a: [3,142,0,2,1,94]
s: 142
i: [1,96,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,147,0,2,1,96]
s: 147
i: [1,95,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,152,0,2,1,95]
s: 152
i: [1,67,0,9,3,1,0,0] t3_subset_1
s: 67
i: [1,68,0,9,3,1,0,0] t3_subset_1
s: 68
i: [1,73,0,3,1,0,0,1] fc1_subset_1
s: 73
i: [1,89,0,3,1,0,0,1] rc2_subset_1
s: 89
i: [1,90,0,3,1,0,0,255] reflexivity_r1_tarski
s: 90
i: [1,97,0,3,1,0,0,255] d10_xboole_0
s: 97
i: [1,98,0,3,1,0,0,255] d10_xboole_0
s: 98
i: [1,72,0,3,1,0,0,255] d2_xboole_0
a: [3,127,0,3,1,72]
s: 127
r: 117
i: [1,74,0,4,2,0,0,255] cc2_zfmisc_1
s: 74
i: [1,77,0,4,2,0,0,255] cc1_zfmisc_1
s: 77
i: [1,87,0,4,1,0,0,1] existence_m1_subset_1
s: 87
d: [2,155,1,2,1,1,40,74,122]
a: [3,161,1,2,1,155]
s: 161
i: [1,76,0,5,2,0,0,1] rc1_subset_1
s: 76
d: [2,157,1,2,1,1,40,74,142]
a: [3,166,1,2,1,157]
s: 166
i: [1,78,0,5,2,0,0,255] t6_boole
s: 78
i: [1,86,0,5,2,0,0,1] t7_boole
s: 86
d: [2,156,1,3,1,0,40,74,89]
s: 156
i: [1,88,0,5,1,0,0,1] rc2_subset_1
s: 88
i: [1,91,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 91
i: [1,75,0,7,2,0,0,1] rc1_subset_1
s: 75
i: [1,85,0,7,3,0,0,1] t8_boole
s: 85
d: [2,168,1,7,3,0,44,78,78]
s: 168
i: [1,84,0,9,3,0,0,255] d10_xboole_0
s: 84
i: [1,79,0,17,4,0,0,255] t2_subset_1
s: 79
i: [1,80,0,17,4,0,0,255] t2_subset_1
s: 80
i: [1,81,0,17,4,0,0,255] t2_subset_1
s: 81
d: [2,172,1,23,5,0,40,80,67]
s: 172
d: [2,173,1,23,5,0,40,80,68]
s: 173
d: [2,178,2,28,7,0,40,172,79]
d: [2,179,2,17,4,0,18,178]
d: [2,180,2,13,3,1,21,179,112]
s: 180
d: [2,175,1,13,4,0,40,80,86]
s: 175
d: [2,174,1,17,4,0,40,80,91]
s: 174
d: [2,187,3,9,3,1,40,180,86]
a: [3,199,3,2,1,187]
s: 199
d: [2,204,2,9,3,1,40,175,107]
s: 204
d: [2,184,3,18,5,1,40,180,81]
d: [2,188,3,11,3,1,18,184]
d: [2,189,3,7,2,2,21,188,107]
d: [2,190,3,3,1,2,21,189,112]
a: [3,194,3,3,1,190]
s: 194
d: [2,205,2,9,3,1,40,175,112]
s: 205
d: [2,181,2,28,7,0,40,173,79]
d: [2,182,2,17,4,0,18,181]
d: [2,183,2,13,3,1,21,182,107]
s: 183
d: [2,236,4,6,2,1,40,194,84]
d: [2,237,4,3,1,2,21,236,102]
d: [2,260,3,18,5,1,40,183,81]
d: [2,264,3,11,3,1,18,260]
d: [2,265,3,7,2,2,21,264,112]
d: [2,266,3,3,1,2,21,265,107]
f: 266,237,190,65,66,69
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t3_subset_1
% SZS output start Proof for subset_1__t3_subset_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => (! [X3] : (m1_subset_1(X3,X0) => (r2_hidden(X3,X1) <=> r2_hidden(X3,X2))) => X1 = X2))) [input t3_subset_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => (! [X3] : (m1_subset_1(X3,X0) => (r2_hidden(X3,X1) <=> r2_hidden(X3,X2))) => X1 = X2))) [negated conjecture 1]
6. ! [X0] : ! [X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [input d10_xboole_0]
22. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => (! [X3] : (m1_subset_1(X3,X0) => (r2_hidden(X3,X1) => r2_hidden(X3,X2))) => r1_tarski(X1,X2)))) [input t2_subset_1]
26. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => (! [X3] : (m1_subset_1(X3,X0) => (r2_hidden(X3,X1) <=> r2_hidden(X3,X2))) => X1 = X2))) [flattening 2]
27. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => (! [X3] : (m1_subset_1(X3,X0) => (r2_hidden(X3,X1) => r2_hidden(X3,X2))) => r1_tarski(X1,X2)))) [flattening 22]
28. ! [X0,X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [flattening 6]
32. ? [X0,X1] : (? [X2] : ((X1 != X2 & ! [X3] : ((r2_hidden(X3,X1) <=> r2_hidden(X3,X2)) | ~m1_subset_1(X3,X0))) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 26]
33. ? [X0,X1] : (? [X2] : (X1 != X2 & ! [X3] : ((r2_hidden(X3,X1) <=> r2_hidden(X3,X2)) | ~m1_subset_1(X3,X0)) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 32]
38. ! [X0,X1] : (! [X2] : ((r1_tarski(X1,X2) | ? [X3] : ((~r2_hidden(X3,X2) & r2_hidden(X3,X1)) & m1_subset_1(X3,X0))) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 27]
39. ! [X0,X1] : (! [X2] : (r1_tarski(X1,X2) | ? [X3] : (~r2_hidden(X3,X2) & r2_hidden(X3,X1) & m1_subset_1(X3,X0)) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 38]
43. ? [X0,X1] : (? [X2] : (X1 != X2 & ! [X3] : (((r2_hidden(X3,X1) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X2) | ~r2_hidden(X3,X1))) | ~m1_subset_1(X3,X0)) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [nnf transformation 33]
44. ? [X0,X1] : (? [X2] : (X1 != X2 & ! [X3] : (((r2_hidden(X3,X1) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X2) | ~r2_hidden(X3,X1))) | ~m1_subset_1(X3,X0)) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (? [X2] : (sK1 != X2 & ! [X3] : (((r2_hidden(X3,sK1) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X2) | ~r2_hidden(X3,sK1))) | ~m1_subset_1(X3,sK0)) & m1_subset_1(X2,k1_zfmisc_1(sK0))) & m1_subset_1(sK1,k1_zfmisc_1(sK0))) [choice axiom]
45. ? [X2] : (sK1 != X2 & ! [X3] : (((r2_hidden(X3,sK1) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X2) | ~r2_hidden(X3,sK1))) | ~m1_subset_1(X3,sK0)) & m1_subset_1(X2,k1_zfmisc_1(sK0))) => (sK1 != sK2 & ! [X3] : (((r2_hidden(X3,sK1) | ~r2_hidden(X3,sK2)) & (r2_hidden(X3,sK2) | ~r2_hidden(X3,sK1))) | ~m1_subset_1(X3,sK0)) & m1_subset_1(sK2,k1_zfmisc_1(sK0))) [choice axiom]
46. (sK1 != sK2 & ! [X3] : (((r2_hidden(X3,sK1) | ~r2_hidden(X3,sK2)) & (r2_hidden(X3,sK2) | ~r2_hidden(X3,sK1))) | ~m1_subset_1(X3,sK0)) & m1_subset_1(sK2,k1_zfmisc_1(sK0))) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) [skolemisation 43,45,44]
49. ! [X2,X1,X0] : (? [X3] : (~r2_hidden(X3,X2) & r2_hidden(X3,X1) & m1_subset_1(X3,X0)) => (~r2_hidden(sK4(X0,X1,X2),X2) & r2_hidden(sK4(X0,X1,X2),X1) & m1_subset_1(sK4(X0,X1,X2),X0))) [choice axiom]
50. ! [X0,X1] : (! [X2] : (r1_tarski(X1,X2) | (~r2_hidden(sK4(X0,X1,X2),X2) & r2_hidden(sK4(X0,X1,X2),X1) & m1_subset_1(sK4(X0,X1,X2),X0)) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [skolemisation 39,49]
51. ! [X0,X1] : ((X0 = X1 | (~r1_tarski(X1,X0) | ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [nnf transformation 28]
52. ! [X0,X1] : ((X0 = X1 | ~r1_tarski(X1,X0) | ~r1_tarski(X0,X1)) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [flattening 51]
65. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 46] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
66. m1_subset_1(sK2,k1_zfmisc_1(sK0)) [cnf transformation 46] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
67. ~r2_hidden(X3,sK1) | r2_hidden(X3,sK2) | ~m1_subset_1(X3,sK0) [cnf transformation 46] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
68. ~r2_hidden(X3,sK2) | r2_hidden(X3,sK1) | ~m1_subset_1(X3,sK0) [cnf transformation 46] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
69. sK1 != sK2 [cnf transformation 46] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
79. m1_subset_1(sK4(X0,X1,X2),X0) | r1_tarski(X1,X2) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 50] {a:0,w:17,wCS:6800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. r2_hidden(sK4(X0,X1,X2),X1) | r1_tarski(X1,X2) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 50] {a:0,w:17,wCS:6800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
81. ~r2_hidden(sK4(X0,X1,X2),X2) | r1_tarski(X1,X2) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 50] {a:0,w:17,wCS:6800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. ~r1_tarski(X1,X0) | X0 = X1 | ~r1_tarski(X0,X1) [cnf transformation 52] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
100. 1 <=> sK1 = sK2 [avatar definition]
102. sK1 != sK2 <- (~1) [avatar component clause 100] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
103. ~1 [avatar split clause 69,100]
105. 2 <=> m1_subset_1(sK2,k1_zfmisc_1(sK0)) [avatar definition]
107. m1_subset_1(sK2,k1_zfmisc_1(sK0)) <- (2) [avatar component clause 105] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
108. 2 [avatar split clause 66,105]
110. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
112. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (3) [avatar component clause 110] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. 3 [avatar split clause 65,110]
172. ~m1_subset_1(sK4(X1,sK1,X0),sK0) | ~m1_subset_1(X0,k1_zfmisc_1(X1)) | ~m1_subset_1(sK1,k1_zfmisc_1(X1)) | r2_hidden(sK4(X1,sK1,X0),sK2) | r1_tarski(sK1,X0) [resolution 80,67] {a:1,w:23,wCS:2300,nSel:1,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
173. ~m1_subset_1(sK4(X3,sK2,X2),sK0) | ~m1_subset_1(X2,k1_zfmisc_1(X3)) | ~m1_subset_1(sK2,k1_zfmisc_1(X3)) | r2_hidden(sK4(X3,sK2,X2),sK1) | r1_tarski(sK2,X2) [resolution 80,68] {a:1,w:23,wCS:2300,nSel:1,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
178. ~m1_subset_1(X0,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) | r2_hidden(sK4(sK0,sK1,X0),sK2) | r1_tarski(sK1,X0) | r1_tarski(sK1,X0) | ~m1_subset_1(X0,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) [resolution 172,79] {a:2,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
179. ~m1_subset_1(X0,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) | r2_hidden(sK4(sK0,sK1,X0),sK2) | r1_tarski(sK1,X0) [duplicate literal removal 178] {a:2,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
180. r2_hidden(sK4(sK0,sK1,X0),sK2) | ~m1_subset_1(X0,k1_zfmisc_1(sK0)) | r1_tarski(sK1,X0) <- (3) [subsumption resolution 179,112] {a:2,w:13,wCS:1300,nSel:1,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
181. ~m1_subset_1(X0,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | r2_hidden(sK4(sK0,sK2,X0),sK1) | r1_tarski(sK2,X0) | r1_tarski(sK2,X0) | ~m1_subset_1(X0,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) [resolution 173,79] {a:2,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
182. ~m1_subset_1(X0,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | r2_hidden(sK4(sK0,sK2,X0),sK1) | r1_tarski(sK2,X0) [duplicate literal removal 181] {a:2,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
183. r2_hidden(sK4(sK0,sK2,X0),sK1) | ~m1_subset_1(X0,k1_zfmisc_1(sK0)) | r1_tarski(sK2,X0) <- (2) [subsumption resolution 182,107] {a:2,w:13,wCS:1300,nSel:1,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
184. ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | r1_tarski(sK1,sK2) | r1_tarski(sK1,sK2) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (3) [resolution 180,81] {a:3,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:9,thDist:-9}
188. ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | r1_tarski(sK1,sK2) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (3) [duplicate literal removal 184] {a:3,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:9,thDist:-9}
189. r1_tarski(sK1,sK2) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (2, 3) [subsumption resolution 188,107] {a:3,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:9,thDist:-9}
190. r1_tarski(sK1,sK2) <- (2, 3) [subsumption resolution 189,112] {a:3,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:9,thDist:-9}
192. 14 <=> r1_tarski(sK1,sK2) [avatar definition]
194. r1_tarski(sK1,sK2) <- (14) [avatar component clause 192] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
195. 14 | ~2 | ~3 [avatar split clause 190,110,105,192]
236. sK1 = sK2 | ~r1_tarski(sK2,sK1) <- (14) [resolution 194,84] {a:4,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
237. ~r1_tarski(sK2,sK1) <- (~1, 14) [subsumption resolution 236,102] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
239. 21 <=> r1_tarski(sK2,sK1) [avatar definition]
242. ~21 | 1 | ~14 [avatar split clause 237,192,100,239]
260. ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) | r1_tarski(sK2,sK1) | r1_tarski(sK2,sK1) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) <- (2) [resolution 183,81] {a:3,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:9,thDist:-9}
264. ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) | r1_tarski(sK2,sK1) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) <- (2) [duplicate literal removal 260] {a:3,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:9,thDist:-9}
265. r1_tarski(sK2,sK1) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) <- (2, 3) [subsumption resolution 264,112] {a:3,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:9,thDist:-9}
266. r1_tarski(sK2,sK1) <- (2, 3) [subsumption resolution 265,107] {a:3,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:9,thDist:-9}
267. 21 | ~2 | ~3 [avatar split clause 266,110,105,239]
272. $false [avatar sat refutation 103,108,113,195,242,267] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:156,thDist:-156}
% SZS output end Proof for subset_1__t3_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% Introduced skolems: 11
% 
% >>> Saturation
% Initial clauses: 32
% Generated clauses: 107
% Active clauses: 42
% Passive clauses: 60
% Final active clauses: 41
% Final passive clauses: 14
% Main loop iterations started: 46
% 
% >>> Simplifying Inferences
% Duplicate literals: 14
% Fw subsumption resolutions: 9
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Equational tautologies: 1
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 38
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 25
% SAT solver unit clauses: 11
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

