i: [1,283,0,2,1,1,0,0] t27_cfuncdom
a: [3,519,0,2,1,283]
s: 519
i: [1,284,0,42,14,1,0,0] t27_cfuncdom
a: [3,466,0,3,1,284]
s: 466
i: [1,285,0,2,1,0,0,255] fc1_numbers
a: [3,524,0,2,1,285]
s: 524
i: [1,286,0,2,1,0,0,255] fc2_numbers
a: [3,529,0,2,1,286]
s: 529
i: [1,287,0,2,1,0,0,255] fc1_xboole_0
a: [3,534,0,2,1,287]
s: 534
i: [1,415,0,2,1,0,0,1] rc2_xboole_0
a: [3,564,0,2,1,415]
s: 564
i: [1,416,0,2,1,0,0,255] existence_l5_algstr_0
a: [3,569,0,2,1,416]
s: 569
i: [1,417,0,2,1,0,0,255] existence_l4_struct_0
a: [3,574,0,2,1,417]
s: 574
i: [1,418,0,2,1,0,0,255] existence_l1_algstr_0
a: [3,579,0,2,1,418]
s: 579
i: [1,419,0,2,1,0,0,255] existence_l6_algstr_0
a: [3,584,0,2,1,419]
s: 584
i: [1,420,0,2,1,0,0,255] existence_l4_algstr_0
a: [3,589,0,2,1,420]
s: 589
i: [1,421,0,2,1,0,0,4] existence_l3_struct_0
a: [3,594,0,2,1,421]
s: 594
i: [1,422,0,2,1,0,0,2] existence_l3_algstr_0
a: [3,599,0,2,1,422]
s: 599
i: [1,423,0,2,1,0,0,4] existence_l2_struct_0
a: [3,604,0,2,1,423]
s: 604
i: [1,424,0,2,1,0,0,255] existence_l2_algstr_0
a: [3,609,0,2,1,424]
s: 609
i: [1,425,0,2,1,0,0,255] existence_l1_clvect_1
a: [3,614,0,2,1,425]
s: 614
i: [1,426,0,2,1,0,0,1] existence_l1_cfuncdom
a: [3,619,0,2,1,426]
s: 619
i: [1,427,0,2,1,0,0,1] existence_l1_cfuncdom
a: [3,624,0,2,1,427]
s: 624
i: [1,428,0,2,1,0,0,1] rc1_xboole_0
a: [3,629,0,2,1,428]
s: 629
i: [1,429,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,634,0,2,1,429]
s: 634
i: [1,433,0,2,1,0,0,255] rc4_vectsp_1
a: [3,639,0,2,1,433]
s: 639
i: [1,432,0,2,1,0,0,255] rc4_vectsp_1
a: [3,644,0,2,1,432]
s: 644
i: [1,431,0,2,1,0,0,255] rc4_vectsp_1
a: [3,649,0,2,1,431]
s: 649
i: [1,430,0,2,1,0,0,255] rc4_vectsp_1
a: [3,654,0,2,1,430]
s: 654
i: [1,437,0,2,1,0,0,1] rc3_vectsp_1
a: [3,659,0,2,1,437]
s: 659
i: [1,436,0,2,1,0,0,1] rc3_vectsp_1
a: [3,664,0,2,1,436]
s: 664
i: [1,435,0,2,1,0,0,1] rc3_vectsp_1
a: [3,669,0,2,1,435]
s: 669
i: [1,434,0,2,1,0,0,1] rc3_vectsp_1
a: [3,674,0,2,1,434]
s: 674
i: [1,439,0,2,1,0,0,255] rc2_clvect_1
a: [3,679,0,2,1,439]
s: 679
i: [1,438,0,2,1,0,0,255] rc2_clvect_1
a: [3,684,0,2,1,438]
s: 684
i: [1,441,0,2,1,0,0,1] rc2_cfuncdom
a: [3,689,0,2,1,441]
s: 689
i: [1,440,0,2,1,0,0,1] rc2_cfuncdom
a: [3,694,0,2,1,440]
s: 694
i: [1,456,0,2,1,0,0,1] rc3_cfuncdom
a: [3,699,0,2,1,456]
s: 699
i: [1,455,0,2,1,0,0,1] rc3_cfuncdom
a: [3,704,0,2,1,455]
s: 704
i: [1,454,0,2,1,0,0,1] rc3_cfuncdom
a: [3,709,0,2,1,454]
s: 709
i: [1,453,0,2,1,0,0,1] rc3_cfuncdom
a: [3,714,0,2,1,453]
s: 714
i: [1,452,0,2,1,0,0,1] rc3_cfuncdom
a: [3,719,0,2,1,452]
s: 719
i: [1,451,0,2,1,0,0,1] rc3_cfuncdom
a: [3,724,0,2,1,451]
s: 724
i: [1,450,0,2,1,0,0,1] rc3_cfuncdom
a: [3,729,0,2,1,450]
s: 729
i: [1,449,0,2,1,0,0,1] rc3_cfuncdom
a: [3,734,0,2,1,449]
s: 734
i: [1,448,0,2,1,0,0,1] rc3_cfuncdom
a: [3,739,0,2,1,448]
s: 739
i: [1,447,0,2,1,0,0,1] rc3_cfuncdom
a: [3,744,0,2,1,447]
s: 744
i: [1,446,0,2,1,0,0,1] rc3_cfuncdom
a: [3,749,0,2,1,446]
s: 749
i: [1,445,0,2,1,0,0,1] rc3_cfuncdom
a: [3,754,0,2,1,445]
s: 754
i: [1,444,0,2,1,0,0,1] rc3_cfuncdom
a: [3,759,0,2,1,444]
s: 759
i: [1,443,0,2,1,0,0,1] rc3_cfuncdom
a: [3,764,0,2,1,443]
s: 764
i: [1,442,0,2,1,0,0,1] rc3_cfuncdom
a: [3,769,0,2,1,442]
s: 769
i: [1,458,0,2,1,0,0,2] rc1_cfuncdom
a: [3,774,0,2,1,458]
s: 774
i: [1,457,0,2,1,0,0,2] rc1_cfuncdom
a: [3,779,0,2,1,457]
s: 779
i: [1,375,0,3,1,0,0,255] reflexivity_r1_tarski
s: 375
i: [1,288,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,539,0,3,1,288]
s: 539
i: [1,289,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,544,0,3,1,289]
s: 544
r: 534
d: [2,781,0,2,1,2,23,534,544]
a: [3,785,0,2,1,781]
s: 785
i: [1,290,0,3,1,0,0,255] dt_k6_complex1
a: [3,549,0,3,1,290]
s: 549
i: [1,338,0,4,2,0,0,255] dt_l5_algstr_0
s: 338
i: [1,339,0,4,2,0,0,255] dt_l5_algstr_0
s: 339
d: [2,787,1,2,1,1,40,338,569]
a: [3,791,1,2,1,787]
s: 791
i: [1,340,0,4,2,0,0,255] dt_l4_struct_0
s: 340
d: [2,793,1,2,1,1,40,339,569]
a: [3,797,1,2,1,793]
s: 797
i: [1,341,0,4,2,0,0,255] dt_l4_struct_0
s: 341
i: [1,345,0,4,2,0,0,255] dt_l6_algstr_0
s: 345
d: [2,799,1,2,1,1,40,340,574]
a: [3,803,1,2,1,799]
s: 803
i: [1,346,0,4,2,0,0,255] dt_l6_algstr_0
s: 346
d: [2,811,1,2,1,1,40,341,574]
a: [3,816,1,2,1,811]
s: 816
i: [1,347,0,4,2,0,0,255] dt_l4_algstr_0
s: 347
i: [1,348,0,4,2,0,0,255] dt_l4_algstr_0
s: 348
d: [2,823,1,2,1,1,40,345,584]
a: [3,827,1,2,1,823]
s: 827
i: [1,351,0,4,2,0,0,255] dt_l2_algstr_0
s: 351
d: [2,829,1,2,1,1,40,346,584]
a: [3,833,1,2,1,829]
s: 833
i: [1,352,0,4,2,0,0,255] dt_l2_algstr_0
s: 352
i: [1,356,0,4,2,0,0,255] dt_l1_clvect_1
s: 356
d: [2,835,1,2,1,1,40,347,589]
a: [3,841,1,2,1,835]
s: 841
i: [1,360,0,4,2,0,0,255] dt_l1_cfuncdom
s: 360
d: [2,836,1,2,1,1,40,347,654]
a: [3,846,1,2,1,836]
s: 846
i: [1,361,0,4,2,0,0,255] dt_l1_cfuncdom
s: 361
i: [1,362,0,4,2,0,0,255] dt_l1_cfuncdom
s: 362
d: [2,849,1,2,1,1,40,348,589]
a: [3,855,1,2,1,849]
s: 855
i: [1,363,0,4,2,0,0,255] dt_l1_cfuncdom
s: 363
d: [2,850,1,2,1,1,40,348,654]
a: [3,860,1,2,1,850]
s: 860
i: [1,374,0,4,1,0,0,255] existence_m1_subset_1
s: 374
i: [1,388,0,4,1,0,0,255] rc1_relset_1
s: 388
d: [2,867,1,2,1,1,40,351,609]
a: [3,872,1,2,1,867]
s: 872
i: [1,291,0,4,1,0,0,255] dt_k5_numbers
a: [3,554,0,4,1,291]
s: 554
d: [2,891,1,2,1,1,40,352,609]
a: [3,896,1,2,1,891]
s: 896
i: [1,292,0,4,1,0,0,255] dt_k6_numbers
a: [3,559,0,4,1,292]
s: 559
i: [1,300,0,5,2,0,0,1] fc4_cfuncdom
s: 300
d: [2,903,1,2,1,1,40,356,614]
a: [3,908,1,2,1,903]
s: 908
i: [1,301,0,5,2,0,0,1] fc4_cfuncdom
s: 301
d: [2,904,1,2,1,1,40,356,684]
a: [3,913,1,2,1,904]
s: 913
i: [1,302,0,5,2,0,0,255] fc5_cfuncdom
s: 302
i: [1,303,0,5,2,0,0,1] fc6_cfuncdom
s: 303
d: [2,999,1,0,0,2,52,519,466,303]
e: 999
d: [2,1000,1,2,1,1,40,303,466]
d: [2,1002,1,0,0,2,21,1000,519]
e: 1002
r: 466
a: [3,465,0,3,1,284]
s: 465
a: [3,470,0,3,1,284]
s: 470
d: [2,915,1,2,1,1,40,360,619]
a: [3,923,1,2,1,915]
s: 923
i: [1,304,0,5,2,0,0,1] fc6_cfuncdom
s: 304
d: [2,1016,1,0,0,2,52,519,470,304]
e: 1016
d: [2,1017,1,2,1,1,40,304,470]
d: [2,1019,1,0,0,2,21,1017,519]
e: 1019
r: 470
a: [3,469,0,3,1,284]
s: 469
a: [3,474,0,3,1,284]
s: 474
i: [1,305,0,5,2,0,0,1] fc6_cfuncdom
s: 305
d: [2,1021,1,0,0,2,52,519,474,305]
e: 1021
d: [2,1022,1,2,1,1,40,305,474]
d: [2,1024,1,0,0,2,21,1022,519]
e: 1024
r: 474
a: [3,473,0,3,1,284]
s: 473
a: [3,478,0,3,1,284]
s: 478
i: [1,306,0,5,2,0,0,1] fc6_cfuncdom
s: 306
d: [2,1026,1,0,0,2,52,519,478,306]
e: 1026
d: [2,1027,1,2,1,1,40,306,478]
d: [2,1029,1,0,0,2,21,1027,519]
e: 1029
r: 478
a: [3,477,0,3,1,284]
s: 477
a: [3,482,0,3,1,284]
s: 482
d: [2,916,1,2,1,1,40,360,624]
a: [3,928,1,2,1,916]
s: 928
i: [1,307,0,5,2,0,0,1] fc6_cfuncdom
s: 307
d: [2,1043,1,0,0,2,52,519,482,307]
e: 1043
d: [2,1044,1,2,1,1,40,307,482]
d: [2,1046,1,0,0,2,21,1044,519]
e: 1046
r: 482
a: [3,481,0,3,1,284]
s: 481
a: [3,486,0,3,1,284]
s: 486
i: [1,308,0,5,2,0,0,1] fc6_cfuncdom
s: 308
d: [2,1048,1,0,0,2,52,519,486,308]
e: 1048
d: [2,1049,1,2,1,1,40,308,486]
d: [2,1051,1,0,0,2,21,1049,519]
e: 1051
r: 486
a: [3,485,0,3,1,284]
s: 485
a: [3,490,0,3,1,284]
s: 490
i: [1,309,0,5,2,0,0,1] fc6_cfuncdom
s: 309
d: [2,1053,1,0,0,2,52,519,490,309]
e: 1053
d: [2,1054,1,2,1,1,40,309,490]
d: [2,1056,1,0,0,2,21,1054,519]
e: 1056
r: 490
a: [3,489,0,3,1,284]
s: 489
a: [3,494,0,3,1,284]
s: 494
d: [2,917,1,2,1,1,40,360,694]
a: [3,933,1,2,1,917]
s: 933
i: [1,310,0,5,2,0,0,1] fc6_cfuncdom
s: 310
d: [2,1070,1,0,0,2,52,519,494,310]
e: 1070
d: [2,1071,1,2,1,1,40,310,494]
d: [2,1073,1,0,0,2,21,1071,519]
e: 1073
r: 494
a: [3,493,0,3,1,284]
s: 493
a: [3,498,0,3,1,284]
s: 498
i: [1,311,0,5,2,0,0,1] fc6_cfuncdom
s: 311
d: [2,1075,1,0,0,2,52,519,498,311]
e: 1075
d: [2,1076,1,2,1,1,40,311,498]
d: [2,1078,1,0,0,2,21,1076,519]
e: 1078
r: 498
a: [3,497,0,3,1,284]
s: 497
a: [3,502,0,3,1,284]
s: 502
i: [1,312,0,5,2,0,0,1] fc6_cfuncdom
s: 312
d: [2,1080,1,0,0,2,52,519,502,312]
e: 1080
d: [2,1081,1,2,1,1,40,312,502]
d: [2,1083,1,0,0,2,21,1081,519]
e: 1083
r: 502
a: [3,501,0,3,1,284]
s: 501
a: [3,506,0,3,1,284]
s: 506
d: [2,918,1,2,1,1,40,360,769]
a: [3,938,1,2,1,918]
s: 938
i: [1,313,0,5,2,0,0,1] fc6_cfuncdom
s: 313
d: [2,1097,1,0,0,2,52,519,506,313]
e: 1097
d: [2,1098,1,2,1,1,40,313,506]
d: [2,1100,1,0,0,2,21,1098,519]
e: 1100
r: 506
a: [3,505,0,3,1,284]
s: 505
a: [3,510,0,3,1,284]
s: 510
i: [1,314,0,5,2,0,0,1] fc6_cfuncdom
s: 314
d: [2,919,1,2,1,1,40,360,779]
a: [3,943,1,2,1,919]
s: 943
i: [1,315,0,5,2,0,0,1] fc6_cfuncdom
s: 315
d: [2,1114,1,0,0,2,52,519,510,315]
e: 1114
d: [2,1115,1,2,1,1,40,315,510]
d: [2,1117,1,0,0,2,21,1115,519]
e: 1117
r: 510
a: [3,509,0,3,1,284]
s: 509
a: [3,514,0,3,1,284]
s: 514
i: [1,316,0,5,2,0,0,1] dt_k8_cfuncdom
s: 316
d: [2,945,1,2,1,1,40,361,619]
a: [3,953,1,2,1,945]
s: 953
i: [1,317,0,5,2,0,0,1] dt_k8_cfuncdom
s: 317
d: [2,1121,1,0,0,2,52,519,514,317]
e: 1121
d: [2,1122,1,2,1,1,40,317,514]
d: [2,1126,1,0,0,2,21,1122,519]
e: 1126
r: 514
a: [3,462,0,3,1,284]
s: 462
d: [2,1128,1,0,0,2,52,519,462,300]
e: 1128
d: [2,1129,1,2,1,1,40,462,300]
d: [2,1131,1,0,0,2,21,1129,519]
e: 1131
f: 1131,1126,1117,1100,1083,1078,1073,1056,1051,1046,1029,1024,1019,1002,283,284
% Refutation found. Thanks to Tanya!
% SZS status Theorem for cfuncdom__t27_cfuncdom
% SZS output start Proof for cfuncdom__t27_cfuncdom
1. ! [X0] : (~v1_xboole_0(X0) => (l1_cfuncdom(k8_cfuncdom(X0)) & v2_cfuncdom(k8_cfuncdom(X0)) & v3_vectsp_1(k8_cfuncdom(X0)) & v1_vectsp_1(k8_cfuncdom(X0)) & v5_group_1(k8_cfuncdom(X0)) & v3_group_1(k8_cfuncdom(X0)) & v4_clvect_1(k8_cfuncdom(X0)) & v3_clvect_1(k8_cfuncdom(X0)) & v2_clvect_1(k8_cfuncdom(X0)) & v4_rlvect_1(k8_cfuncdom(X0)) & v3_rlvect_1(k8_cfuncdom(X0)) & v2_rlvect_1(k8_cfuncdom(X0)) & v13_algstr_0(k8_cfuncdom(X0)) & ~v2_struct_0(k8_cfuncdom(X0)))) [input t27_cfuncdom]
2. ~! [X0] : (~v1_xboole_0(X0) => (l1_cfuncdom(k8_cfuncdom(X0)) & v2_cfuncdom(k8_cfuncdom(X0)) & v3_vectsp_1(k8_cfuncdom(X0)) & v1_vectsp_1(k8_cfuncdom(X0)) & v5_group_1(k8_cfuncdom(X0)) & v3_group_1(k8_cfuncdom(X0)) & v4_clvect_1(k8_cfuncdom(X0)) & v3_clvect_1(k8_cfuncdom(X0)) & v2_clvect_1(k8_cfuncdom(X0)) & v4_rlvect_1(k8_cfuncdom(X0)) & v3_rlvect_1(k8_cfuncdom(X0)) & v2_rlvect_1(k8_cfuncdom(X0)) & v13_algstr_0(k8_cfuncdom(X0)) & ~v2_struct_0(k8_cfuncdom(X0)))) [negated conjecture 1]
107. ! [X0] : (~v1_xboole_0(X0) => (l1_cfuncdom(k8_cfuncdom(X0)) & v1_cfuncdom(k8_cfuncdom(X0)))) [input dt_k8_cfuncdom]
110. ! [X0] : (~v1_xboole_0(X0) => (v1_cfuncdom(k8_cfuncdom(X0)) & ~v2_struct_0(k8_cfuncdom(X0)))) [input fc4_cfuncdom]
112. ! [X0] : (~v1_xboole_0(X0) => (v2_cfuncdom(k8_cfuncdom(X0)) & v1_cfuncdom(k8_cfuncdom(X0)) & v3_vectsp_1(k8_cfuncdom(X0)) & v1_vectsp_1(k8_cfuncdom(X0)) & v5_group_1(k8_cfuncdom(X0)) & v3_group_1(k8_cfuncdom(X0)) & v4_clvect_1(k8_cfuncdom(X0)) & v3_clvect_1(k8_cfuncdom(X0)) & v2_clvect_1(k8_cfuncdom(X0)) & v4_rlvect_1(k8_cfuncdom(X0)) & v3_rlvect_1(k8_cfuncdom(X0)) & v2_rlvect_1(k8_cfuncdom(X0)) & v13_algstr_0(k8_cfuncdom(X0)))) [input fc6_cfuncdom]
151. ? [X0] : ((~l1_cfuncdom(k8_cfuncdom(X0)) | ~v2_cfuncdom(k8_cfuncdom(X0)) | ~v3_vectsp_1(k8_cfuncdom(X0)) | ~v1_vectsp_1(k8_cfuncdom(X0)) | ~v5_group_1(k8_cfuncdom(X0)) | ~v3_group_1(k8_cfuncdom(X0)) | ~v4_clvect_1(k8_cfuncdom(X0)) | ~v3_clvect_1(k8_cfuncdom(X0)) | ~v2_clvect_1(k8_cfuncdom(X0)) | ~v4_rlvect_1(k8_cfuncdom(X0)) | ~v3_rlvect_1(k8_cfuncdom(X0)) | ~v2_rlvect_1(k8_cfuncdom(X0)) | ~v13_algstr_0(k8_cfuncdom(X0)) | v2_struct_0(k8_cfuncdom(X0))) & ~v1_xboole_0(X0)) [ennf transformation 2]
159. ! [X0] : ((v1_cfuncdom(k8_cfuncdom(X0)) & ~v2_struct_0(k8_cfuncdom(X0))) | v1_xboole_0(X0)) [ennf transformation 110]
161. ! [X0] : ((v2_cfuncdom(k8_cfuncdom(X0)) & v1_cfuncdom(k8_cfuncdom(X0)) & v3_vectsp_1(k8_cfuncdom(X0)) & v1_vectsp_1(k8_cfuncdom(X0)) & v5_group_1(k8_cfuncdom(X0)) & v3_group_1(k8_cfuncdom(X0)) & v4_clvect_1(k8_cfuncdom(X0)) & v3_clvect_1(k8_cfuncdom(X0)) & v2_clvect_1(k8_cfuncdom(X0)) & v4_rlvect_1(k8_cfuncdom(X0)) & v3_rlvect_1(k8_cfuncdom(X0)) & v2_rlvect_1(k8_cfuncdom(X0)) & v13_algstr_0(k8_cfuncdom(X0))) | v1_xboole_0(X0)) [ennf transformation 112]
162. ! [X0] : ((l1_cfuncdom(k8_cfuncdom(X0)) & v1_cfuncdom(k8_cfuncdom(X0))) | v1_xboole_0(X0)) [ennf transformation 107]
227. ? [X0] : ((~l1_cfuncdom(k8_cfuncdom(X0)) | ~v2_cfuncdom(k8_cfuncdom(X0)) | ~v3_vectsp_1(k8_cfuncdom(X0)) | ~v1_vectsp_1(k8_cfuncdom(X0)) | ~v5_group_1(k8_cfuncdom(X0)) | ~v3_group_1(k8_cfuncdom(X0)) | ~v4_clvect_1(k8_cfuncdom(X0)) | ~v3_clvect_1(k8_cfuncdom(X0)) | ~v2_clvect_1(k8_cfuncdom(X0)) | ~v4_rlvect_1(k8_cfuncdom(X0)) | ~v3_rlvect_1(k8_cfuncdom(X0)) | ~v2_rlvect_1(k8_cfuncdom(X0)) | ~v13_algstr_0(k8_cfuncdom(X0)) | v2_struct_0(k8_cfuncdom(X0))) & ~v1_xboole_0(X0)) => ((~l1_cfuncdom(k8_cfuncdom(sK0)) | ~v2_cfuncdom(k8_cfuncdom(sK0)) | ~v3_vectsp_1(k8_cfuncdom(sK0)) | ~v1_vectsp_1(k8_cfuncdom(sK0)) | ~v5_group_1(k8_cfuncdom(sK0)) | ~v3_group_1(k8_cfuncdom(sK0)) | ~v4_clvect_1(k8_cfuncdom(sK0)) | ~v3_clvect_1(k8_cfuncdom(sK0)) | ~v2_clvect_1(k8_cfuncdom(sK0)) | ~v4_rlvect_1(k8_cfuncdom(sK0)) | ~v3_rlvect_1(k8_cfuncdom(sK0)) | ~v2_rlvect_1(k8_cfuncdom(sK0)) | ~v13_algstr_0(k8_cfuncdom(sK0)) | v2_struct_0(k8_cfuncdom(sK0))) & ~v1_xboole_0(sK0)) [choice axiom]
228. (~l1_cfuncdom(k8_cfuncdom(sK0)) | ~v2_cfuncdom(k8_cfuncdom(sK0)) | ~v3_vectsp_1(k8_cfuncdom(sK0)) | ~v1_vectsp_1(k8_cfuncdom(sK0)) | ~v5_group_1(k8_cfuncdom(sK0)) | ~v3_group_1(k8_cfuncdom(sK0)) | ~v4_clvect_1(k8_cfuncdom(sK0)) | ~v3_clvect_1(k8_cfuncdom(sK0)) | ~v2_clvect_1(k8_cfuncdom(sK0)) | ~v4_rlvect_1(k8_cfuncdom(sK0)) | ~v3_rlvect_1(k8_cfuncdom(sK0)) | ~v2_rlvect_1(k8_cfuncdom(sK0)) | ~v13_algstr_0(k8_cfuncdom(sK0)) | v2_struct_0(k8_cfuncdom(sK0))) & ~v1_xboole_0(sK0) [skolemisation 151,227]
283. ~v1_xboole_0(sK0) [cnf transformation 228] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
284. ~l1_cfuncdom(k8_cfuncdom(sK0)) | ~v2_cfuncdom(k8_cfuncdom(sK0)) | ~v3_vectsp_1(k8_cfuncdom(sK0)) | ~v1_vectsp_1(k8_cfuncdom(sK0)) | ~v5_group_1(k8_cfuncdom(sK0)) | ~v3_group_1(k8_cfuncdom(sK0)) | ~v4_clvect_1(k8_cfuncdom(sK0)) | ~v3_clvect_1(k8_cfuncdom(sK0)) | ~v2_clvect_1(k8_cfuncdom(sK0)) | ~v4_rlvect_1(k8_cfuncdom(sK0)) | ~v3_rlvect_1(k8_cfuncdom(sK0)) | ~v2_rlvect_1(k8_cfuncdom(sK0)) | ~v13_algstr_0(k8_cfuncdom(sK0)) | v2_struct_0(k8_cfuncdom(sK0)) [cnf transformation 228] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
300. ~v2_struct_0(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 159] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
303. v13_algstr_0(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
304. v2_rlvect_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
305. v3_rlvect_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
306. v4_rlvect_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
307. v2_clvect_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
308. v3_clvect_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
309. v4_clvect_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
310. v3_group_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
311. v5_group_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
312. v1_vectsp_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
313. v3_vectsp_1(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
315. v2_cfuncdom(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 161] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
317. l1_cfuncdom(k8_cfuncdom(X0)) | v1_xboole_0(X0) [cnf transformation 162] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
460. 1 <=> v2_struct_0(k8_cfuncdom(sK0)) [avatar definition]
462. v2_struct_0(k8_cfuncdom(sK0)) <- (1) [avatar component clause 460] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
464. 2 <=> v13_algstr_0(k8_cfuncdom(sK0)) [avatar definition]
466. ~v13_algstr_0(k8_cfuncdom(sK0)) <- (~2) [avatar component clause 464] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
468. 3 <=> v2_rlvect_1(k8_cfuncdom(sK0)) [avatar definition]
470. ~v2_rlvect_1(k8_cfuncdom(sK0)) <- (~3) [avatar component clause 468] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
472. 4 <=> v3_rlvect_1(k8_cfuncdom(sK0)) [avatar definition]
474. ~v3_rlvect_1(k8_cfuncdom(sK0)) <- (~4) [avatar component clause 472] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
476. 5 <=> v4_rlvect_1(k8_cfuncdom(sK0)) [avatar definition]
478. ~v4_rlvect_1(k8_cfuncdom(sK0)) <- (~5) [avatar component clause 476] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
480. 6 <=> v2_clvect_1(k8_cfuncdom(sK0)) [avatar definition]
482. ~v2_clvect_1(k8_cfuncdom(sK0)) <- (~6) [avatar component clause 480] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
484. 7 <=> v3_clvect_1(k8_cfuncdom(sK0)) [avatar definition]
486. ~v3_clvect_1(k8_cfuncdom(sK0)) <- (~7) [avatar component clause 484] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
488. 8 <=> v4_clvect_1(k8_cfuncdom(sK0)) [avatar definition]
490. ~v4_clvect_1(k8_cfuncdom(sK0)) <- (~8) [avatar component clause 488] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
492. 9 <=> v3_group_1(k8_cfuncdom(sK0)) [avatar definition]
494. ~v3_group_1(k8_cfuncdom(sK0)) <- (~9) [avatar component clause 492] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
496. 10 <=> v5_group_1(k8_cfuncdom(sK0)) [avatar definition]
498. ~v5_group_1(k8_cfuncdom(sK0)) <- (~10) [avatar component clause 496] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
500. 11 <=> v1_vectsp_1(k8_cfuncdom(sK0)) [avatar definition]
502. ~v1_vectsp_1(k8_cfuncdom(sK0)) <- (~11) [avatar component clause 500] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
504. 12 <=> v3_vectsp_1(k8_cfuncdom(sK0)) [avatar definition]
506. ~v3_vectsp_1(k8_cfuncdom(sK0)) <- (~12) [avatar component clause 504] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
508. 13 <=> v2_cfuncdom(k8_cfuncdom(sK0)) [avatar definition]
510. ~v2_cfuncdom(k8_cfuncdom(sK0)) <- (~13) [avatar component clause 508] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
512. 14 <=> l1_cfuncdom(k8_cfuncdom(sK0)) [avatar definition]
514. ~l1_cfuncdom(k8_cfuncdom(sK0)) <- (~14) [avatar component clause 512] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
515. 1 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 [avatar split clause 284,512,508,504,500,496,492,488,484,480,476,472,468,464,460]
517. 15 <=> v1_xboole_0(sK0) [avatar definition]
519. ~v1_xboole_0(sK0) <- (~15) [avatar component clause 517] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
520. ~15 [avatar split clause 283,517]
1000. v1_xboole_0(sK0) <- (~2) [resolution 303,466] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1002. $false <- (~2, ~15) [subsumption resolution 1000,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1003. 2 | 15 [avatar contradiction clause 1002]
1017. v1_xboole_0(sK0) <- (~3) [resolution 304,470] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1019. $false <- (~3, ~15) [subsumption resolution 1017,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1020. 3 | 15 [avatar contradiction clause 1019]
1022. v1_xboole_0(sK0) <- (~4) [resolution 305,474] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1024. $false <- (~4, ~15) [subsumption resolution 1022,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1025. 4 | 15 [avatar contradiction clause 1024]
1027. v1_xboole_0(sK0) <- (~5) [resolution 306,478] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1029. $false <- (~5, ~15) [subsumption resolution 1027,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1030. 5 | 15 [avatar contradiction clause 1029]
1044. v1_xboole_0(sK0) <- (~6) [resolution 307,482] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1046. $false <- (~6, ~15) [subsumption resolution 1044,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1047. 6 | 15 [avatar contradiction clause 1046]
1049. v1_xboole_0(sK0) <- (~7) [resolution 308,486] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1051. $false <- (~7, ~15) [subsumption resolution 1049,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1052. 7 | 15 [avatar contradiction clause 1051]
1054. v1_xboole_0(sK0) <- (~8) [resolution 309,490] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1056. $false <- (~8, ~15) [subsumption resolution 1054,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1057. 8 | 15 [avatar contradiction clause 1056]
1071. v1_xboole_0(sK0) <- (~9) [resolution 310,494] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1073. $false <- (~9, ~15) [subsumption resolution 1071,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1074. 9 | 15 [avatar contradiction clause 1073]
1076. v1_xboole_0(sK0) <- (~10) [resolution 311,498] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1078. $false <- (~10, ~15) [subsumption resolution 1076,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1079. 10 | 15 [avatar contradiction clause 1078]
1081. v1_xboole_0(sK0) <- (~11) [resolution 312,502] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1083. $false <- (~11, ~15) [subsumption resolution 1081,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1084. 11 | 15 [avatar contradiction clause 1083]
1098. v1_xboole_0(sK0) <- (~12) [resolution 313,506] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1100. $false <- (~12, ~15) [subsumption resolution 1098,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1101. 12 | 15 [avatar contradiction clause 1100]
1115. v1_xboole_0(sK0) <- (~13) [resolution 315,510] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1117. $false <- (~13, ~15) [subsumption resolution 1115,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1118. 13 | 15 [avatar contradiction clause 1117]
1122. v1_xboole_0(sK0) <- (~14) [resolution 317,514] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1126. $false <- (~14, ~15) [subsumption resolution 1122,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1127. 14 | 15 [avatar contradiction clause 1126]
1129. v1_xboole_0(sK0) <- (1) [resolution 462,300] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1131. $false <- (1, ~15) [subsumption resolution 1129,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1132. ~1 | 15 [avatar contradiction clause 1131]
1133. $false [avatar sat refutation 515,520,1003,1020,1025,1030,1047,1052,1057,1074,1079,1084,1101,1118,1127,1132] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:541,thDist:-541}
% SZS output end Proof for cfuncdom__t27_cfuncdom
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 117
% >>> Preprocessing
% Introduced skolems: 27
% Pure predicates: 8
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 176
% Generated clauses: 396
% Active clauses: 131
% Passive clauses: 250
% Final active clauses: 117
% Final passive clauses: 114
% Main loop iterations started: 136
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 14
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 63
% Unit resulting resolution: 14
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 14
% Sat splitting refutations: 28
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 132
% SAT solver unit clauses: 53
% SAT solver binary clauses: 75
% 
% Memory used [KB]: 10106
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% term sharing: 0.001 s
% unit resulting resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.004 s ( own 0.002 s ) 

