i: [1,105,0,2,1,1,0,0] t59_qc_lang2
a: [3,169,0,2,1,105]
s: 169
i: [1,109,0,4,1,1,0,0] t59_qc_lang2
a: [3,149,0,4,1,109]
s: 149
i: [1,108,0,4,1,1,0,0] t59_qc_lang2
a: [3,154,0,4,1,108]
s: 154
i: [1,107,0,4,1,1,0,0] t59_qc_lang2
a: [3,159,0,4,1,107]
s: 159
i: [1,106,0,4,1,1,0,0] t59_qc_lang2
a: [3,164,0,4,1,106]
s: 164
i: [1,110,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,174,0,2,1,110]
s: 174
i: [1,133,0,2,1,0,0,1] existence_m1_qc_lang1
a: [3,184,0,2,1,133]
s: 184
i: [1,134,0,2,1,0,0,255] rc4_ordinal1
a: [3,189,0,2,1,134]
s: 189
i: [1,135,0,2,1,0,0,255] rc1_ordinal1
a: [3,194,0,2,1,135]
s: 194
i: [1,139,0,2,1,0,0,255] rc2_ordinal1
a: [3,199,0,2,1,139]
s: 199
i: [1,138,0,2,1,0,0,255] rc2_ordinal1
a: [3,204,0,2,1,138]
s: 204
i: [1,137,0,2,1,0,0,255] rc2_ordinal1
a: [3,209,0,2,1,137]
s: 209
i: [1,136,0,2,1,0,0,255] rc2_ordinal1
a: [3,214,0,2,1,136]
s: 214
i: [1,140,0,2,1,0,0,255] rc8_finseq_1
a: [3,219,0,2,1,140]
s: 219
i: [1,142,0,2,1,0,0,255] rc5_ordinal1
a: [3,224,0,2,1,142]
s: 224
i: [1,141,0,2,1,0,0,255] rc5_ordinal1
a: [3,229,0,2,1,141]
s: 229
i: [1,143,0,2,1,0,0,255] rc4_xreal_0
a: [3,234,0,2,1,143]
s: 234
i: [1,111,0,3,1,0,0,255] d2_xboole_0
a: [3,179,0,3,1,111]
s: 179
r: 174
d: [2,236,0,2,1,2,23,174,179]
a: [3,240,0,2,1,236]
s: 240
i: [1,117,0,4,2,0,0,255] cc6_ordinal1
s: 117
i: [1,118,0,4,2,0,0,255] cc7_ordinal1
s: 118
d: [2,242,1,2,1,1,40,117,189]
a: [3,247,1,2,1,242]
s: 247
i: [1,119,0,4,2,0,0,255] cc3_ordinal1
s: 119
d: [2,243,1,2,1,1,40,117,224]
a: [3,252,1,2,1,243]
s: 252
i: [1,121,0,4,2,0,0,255] cc1_ordinal1
s: 121
i: [1,122,0,4,2,0,0,255] cc1_ordinal1
s: 122
d: [2,254,1,2,1,1,40,118,234]
a: [3,259,1,2,1,254]
s: 259
i: [1,129,0,4,1,0,0,1] existence_m1_subset_1
s: 129
d: [2,255,1,2,1,1,40,118,240]
a: [3,264,1,2,1,255]
s: 264
i: [1,112,0,5,2,0,0,1] dt_k9_qc_lang1
s: 112
i: [1,120,0,5,2,0,0,255] t6_boole
s: 120
d: [2,266,1,2,1,1,40,119,234]
a: [3,271,1,2,1,266]
s: 271
i: [1,128,0,5,2,0,0,255] t7_boole
s: 128
d: [2,267,1,2,1,1,40,119,240]
a: [3,276,1,2,1,267]
s: 276
i: [1,124,0,6,3,0,0,255] cc2_ordinal1
s: 124
i: [1,125,0,6,2,0,0,255] t1_subset
s: 125
d: [2,279,1,2,1,1,40,121,194]
a: [3,290,1,2,1,279]
s: 290
i: [1,130,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 130
d: [2,298,1,2,1,1,40,122,194]
a: [3,309,1,2,1,298]
s: 309
i: [1,123,0,7,3,0,0,255] cc5_ordinal1
s: 123
i: [1,127,0,7,3,0,0,2] t8_boole
s: 127
d: [2,350,1,5,2,1,40,123,159]
a: [3,356,1,3,1,350]
s: 356
i: [1,126,0,8,3,0,0,255] t2_subset
s: 126
d: [2,367,1,7,2,1,40,126,159]
a: [3,377,1,3,1,367]
s: 377
d: [2,384,2,0,0,2,52,169,377,112]
e: 384
d: [2,385,2,0,0,2,52,356,377,119]
e: 385
d: [2,386,2,2,1,1,40,377,112]
d: [2,391,2,0,0,2,21,386,169]
e: 391
d: [2,387,2,3,1,1,40,377,119]
d: [2,393,2,0,0,2,21,387,356]
e: 393
r: 377
i: [1,144,0,14,4,0,0,1] d21_qc_lang2
d: [2,145,0,10,3,0,18,144]
s: 145
i: [1,131,0,14,4,0,0,1] irreflexivity_r3_qc_lang2
s: 131
a: [3,376,1,3,1,367]
s: 376
i: [1,132,0,14,4,0,0,1] reflexivity_r2_qc_lang2
s: 132
a: [3,373,1,4,1,367]
s: 373
i: [1,114,0,18,5,0,0,1] d21_qc_lang2
s: 114
i: [1,113,0,21,6,0,0,1] t58_qc_lang2
s: 113
d: [2,417,1,17,5,1,40,113,149]
d: [2,418,1,13,4,2,21,417,159]
d: [2,419,1,9,3,3,21,418,164]
d: [2,420,1,7,2,4,21,419,169]
a: [3,424,1,3,1,420]
s: 424
r: 373
r: 159
r: 154
r: 149
d: [2,428,0,4,1,2,23,154,424]
a: [3,434,0,4,1,428]
s: 434
d: [2,441,1,0,0,3,52,169,164,434,145]
e: 441
d: [2,442,1,6,2,1,40,434,145]
d: [2,446,1,2,1,2,21,442,164]
d: [2,447,1,0,0,3,21,446,169]
e: 447
d: [2,408,1,14,4,1,40,114,154]
d: [2,409,1,10,3,2,21,408,159]
d: [2,410,1,6,2,3,21,409,164]
d: [2,411,1,4,1,4,21,410,169]
f: 447,428,420,411,105,106,107,108,109
% Refutation found. Thanks to Tanya!
% SZS status Theorem for qc_lang2__t59_qc_lang2
% SZS output start Proof for qc_lang2__t59_qc_lang2
1. ! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m1_subset_1(X1,k9_qc_lang1(X0)) => ! [X2] : (m1_subset_1(X2,k9_qc_lang1(X0)) => ~(r2_qc_lang2(X0,X2,X1) & r3_qc_lang2(X0,X1,X2))))) [input t59_qc_lang2]
2. ~! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m1_subset_1(X1,k9_qc_lang1(X0)) => ! [X2] : (m1_subset_1(X2,k9_qc_lang1(X0)) => ~(r2_qc_lang2(X0,X2,X1) & r3_qc_lang2(X0,X1,X2))))) [negated conjecture 1]
18. ! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m1_subset_1(X1,k9_qc_lang1(X0)) => ! [X2] : (m1_subset_1(X2,k9_qc_lang1(X0)) => (r3_qc_lang2(X0,X1,X2) <=> (X1 != X2 & r2_qc_lang2(X0,X1,X2)))))) [input d21_qc_lang2]
38. ! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m1_subset_1(X1,k9_qc_lang1(X0)) => ! [X2] : (m1_subset_1(X2,k9_qc_lang1(X0)) => ((r2_qc_lang2(X0,X2,X1) & r2_qc_lang2(X0,X1,X2)) => X1 = X2)))) [input t58_qc_lang2]
59. ? [X0] : (? [X1] : (? [X2] : ((r2_qc_lang2(X0,X2,X1) & r3_qc_lang2(X0,X1,X2)) & m1_subset_1(X2,k9_qc_lang1(X0))) & m1_subset_1(X1,k9_qc_lang1(X0))) & m1_qc_lang1(X0)) [ennf transformation 2]
60. ? [X0] : (? [X1] : (? [X2] : (r2_qc_lang2(X0,X2,X1) & r3_qc_lang2(X0,X1,X2) & m1_subset_1(X2,k9_qc_lang1(X0))) & m1_subset_1(X1,k9_qc_lang1(X0))) & m1_qc_lang1(X0)) [flattening 59]
62. ! [X0] : (! [X1] : (! [X2] : ((X1 = X2 | (~r2_qc_lang2(X0,X2,X1) | ~r2_qc_lang2(X0,X1,X2))) | ~m1_subset_1(X2,k9_qc_lang1(X0))) | ~m1_subset_1(X1,k9_qc_lang1(X0))) | ~m1_qc_lang1(X0)) [ennf transformation 38]
63. ! [X0] : (! [X1] : (! [X2] : (X1 = X2 | ~r2_qc_lang2(X0,X2,X1) | ~r2_qc_lang2(X0,X1,X2) | ~m1_subset_1(X2,k9_qc_lang1(X0))) | ~m1_subset_1(X1,k9_qc_lang1(X0))) | ~m1_qc_lang1(X0)) [flattening 62]
64. ! [X0] : (! [X1] : (! [X2] : ((r3_qc_lang2(X0,X1,X2) <=> (X1 != X2 & r2_qc_lang2(X0,X1,X2))) | ~m1_subset_1(X2,k9_qc_lang1(X0))) | ~m1_subset_1(X1,k9_qc_lang1(X0))) | ~m1_qc_lang1(X0)) [ennf transformation 18]
83. ? [X0] : (? [X1] : (? [X2] : (r2_qc_lang2(X0,X2,X1) & r3_qc_lang2(X0,X1,X2) & m1_subset_1(X2,k9_qc_lang1(X0))) & m1_subset_1(X1,k9_qc_lang1(X0))) & m1_qc_lang1(X0)) => (? [X1] : (? [X2] : (r2_qc_lang2(sK0,X2,X1) & r3_qc_lang2(sK0,X1,X2) & m1_subset_1(X2,k9_qc_lang1(sK0))) & m1_subset_1(X1,k9_qc_lang1(sK0))) & m1_qc_lang1(sK0)) [choice axiom]
84. ? [X1] : (? [X2] : (r2_qc_lang2(sK0,X2,X1) & r3_qc_lang2(sK0,X1,X2) & m1_subset_1(X2,k9_qc_lang1(sK0))) & m1_subset_1(X1,k9_qc_lang1(sK0))) => (? [X2] : (r2_qc_lang2(sK0,X2,sK1) & r3_qc_lang2(sK0,sK1,X2) & m1_subset_1(X2,k9_qc_lang1(sK0))) & m1_subset_1(sK1,k9_qc_lang1(sK0))) [choice axiom]
85. ? [X2] : (r2_qc_lang2(sK0,X2,sK1) & r3_qc_lang2(sK0,sK1,X2) & m1_subset_1(X2,k9_qc_lang1(sK0))) => (r2_qc_lang2(sK0,sK2,sK1) & r3_qc_lang2(sK0,sK1,sK2) & m1_subset_1(sK2,k9_qc_lang1(sK0))) [choice axiom]
86. ((r2_qc_lang2(sK0,sK2,sK1) & r3_qc_lang2(sK0,sK1,sK2) & m1_subset_1(sK2,k9_qc_lang1(sK0))) & m1_subset_1(sK1,k9_qc_lang1(sK0))) & m1_qc_lang1(sK0) [skolemisation 60,85,84,83]
87. ! [X0] : (! [X1] : (! [X2] : (((r3_qc_lang2(X0,X1,X2) | (X1 = X2 | ~r2_qc_lang2(X0,X1,X2))) & ((X1 != X2 & r2_qc_lang2(X0,X1,X2)) | ~r3_qc_lang2(X0,X1,X2))) | ~m1_subset_1(X2,k9_qc_lang1(X0))) | ~m1_subset_1(X1,k9_qc_lang1(X0))) | ~m1_qc_lang1(X0)) [nnf transformation 64]
88. ! [X0] : (! [X1] : (! [X2] : (((r3_qc_lang2(X0,X1,X2) | X1 = X2 | ~r2_qc_lang2(X0,X1,X2)) & ((X1 != X2 & r2_qc_lang2(X0,X1,X2)) | ~r3_qc_lang2(X0,X1,X2))) | ~m1_subset_1(X2,k9_qc_lang1(X0))) | ~m1_subset_1(X1,k9_qc_lang1(X0))) | ~m1_qc_lang1(X0)) [flattening 87]
105. m1_qc_lang1(sK0) [cnf transformation 86] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
106. m1_subset_1(sK1,k9_qc_lang1(sK0)) [cnf transformation 86] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
107. m1_subset_1(sK2,k9_qc_lang1(sK0)) [cnf transformation 86] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
108. r3_qc_lang2(sK0,sK1,sK2) [cnf transformation 86] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
109. r2_qc_lang2(sK0,sK2,sK1) [cnf transformation 86] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
113. ~r2_qc_lang2(X0,X2,X1) | X1 = X2 | ~r2_qc_lang2(X0,X1,X2) | ~m1_subset_1(X2,k9_qc_lang1(X0)) | ~m1_subset_1(X1,k9_qc_lang1(X0)) | ~m1_qc_lang1(X0) [cnf transformation 63] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
114. ~r3_qc_lang2(X0,X1,X2) | r2_qc_lang2(X0,X1,X2) | ~m1_subset_1(X2,k9_qc_lang1(X0)) | ~m1_subset_1(X1,k9_qc_lang1(X0)) | ~m1_qc_lang1(X0) [cnf transformation 88] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
115. X1 != X2 | ~r3_qc_lang2(X0,X1,X2) | ~m1_subset_1(X2,k9_qc_lang1(X0)) | ~m1_subset_1(X1,k9_qc_lang1(X0)) | ~m1_qc_lang1(X0) [cnf transformation 88] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
144. ~r3_qc_lang2(X0,X2,X2) | ~m1_subset_1(X2,k9_qc_lang1(X0)) | ~m1_subset_1(X2,k9_qc_lang1(X0)) | ~m1_qc_lang1(X0) [equality resolution 115] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:1,thDist:-1}
145. ~r3_qc_lang2(X0,X2,X2) | ~m1_subset_1(X2,k9_qc_lang1(X0)) | ~m1_qc_lang1(X0) [duplicate literal removal 144] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
147. 1 <=> r2_qc_lang2(sK0,sK2,sK1) [avatar definition]
149. r2_qc_lang2(sK0,sK2,sK1) <- (1) [avatar component clause 147] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
150. 1 [avatar split clause 109,147]
152. 2 <=> r3_qc_lang2(sK0,sK1,sK2) [avatar definition]
154. r3_qc_lang2(sK0,sK1,sK2) <- (2) [avatar component clause 152] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
155. 2 [avatar split clause 108,152]
157. 3 <=> m1_subset_1(sK2,k9_qc_lang1(sK0)) [avatar definition]
159. m1_subset_1(sK2,k9_qc_lang1(sK0)) <- (3) [avatar component clause 157] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
160. 3 [avatar split clause 107,157]
162. 4 <=> m1_subset_1(sK1,k9_qc_lang1(sK0)) [avatar definition]
164. m1_subset_1(sK1,k9_qc_lang1(sK0)) <- (4) [avatar component clause 162] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
165. 4 [avatar split clause 106,162]
167. 5 <=> m1_qc_lang1(sK0) [avatar definition]
169. m1_qc_lang1(sK0) <- (5) [avatar component clause 167] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
170. 5 [avatar split clause 105,167]
408. r2_qc_lang2(sK0,sK1,sK2) | ~m1_subset_1(sK2,k9_qc_lang1(sK0)) | ~m1_subset_1(sK1,k9_qc_lang1(sK0)) | ~m1_qc_lang1(sK0) <- (2) [resolution 114,154] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
409. r2_qc_lang2(sK0,sK1,sK2) | ~m1_subset_1(sK1,k9_qc_lang1(sK0)) | ~m1_qc_lang1(sK0) <- (2, 3) [subsumption resolution 408,159] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
410. r2_qc_lang2(sK0,sK1,sK2) | ~m1_qc_lang1(sK0) <- (2, 3, 4) [subsumption resolution 409,164] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
411. r2_qc_lang2(sK0,sK1,sK2) <- (2, 3, 4, 5) [subsumption resolution 410,169] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
413. 43 <=> r2_qc_lang2(sK0,sK1,sK2) [avatar definition]
416. 43 | ~2 | ~3 | ~4 | ~5 [avatar split clause 411,167,162,157,152,413]
417. sK1 = sK2 | ~r2_qc_lang2(sK0,sK1,sK2) | ~m1_subset_1(sK2,k9_qc_lang1(sK0)) | ~m1_subset_1(sK1,k9_qc_lang1(sK0)) | ~m1_qc_lang1(sK0) <- (1) [resolution 113,149] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
418. sK1 = sK2 | ~r2_qc_lang2(sK0,sK1,sK2) | ~m1_subset_1(sK1,k9_qc_lang1(sK0)) | ~m1_qc_lang1(sK0) <- (1, 3) [subsumption resolution 417,159] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
419. sK1 = sK2 | ~r2_qc_lang2(sK0,sK1,sK2) | ~m1_qc_lang1(sK0) <- (1, 3, 4) [subsumption resolution 418,164] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
420. sK1 = sK2 | ~r2_qc_lang2(sK0,sK1,sK2) <- (1, 3, 4, 5) [subsumption resolution 419,169] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
422. 44 <=> sK1 = sK2 [avatar definition]
424. sK1 = sK2 <- (44) [avatar component clause 422] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
425. ~43 | 44 | ~1 | ~3 | ~4 | ~5 [avatar split clause 420,167,162,157,147,422,413]
428. r3_qc_lang2(sK0,sK1,sK1) <- (2, 44) [backward demodulation 154,424] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
432. 45 <=> r3_qc_lang2(sK0,sK1,sK1) [avatar definition]
434. r3_qc_lang2(sK0,sK1,sK1) <- (45) [avatar component clause 432] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
435. 45 | ~2 | ~44 [avatar split clause 428,422,152,432]
442. ~m1_subset_1(sK1,k9_qc_lang1(sK0)) | ~m1_qc_lang1(sK0) <- (45) [resolution 434,145] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
446. ~m1_qc_lang1(sK0) <- (4, 45) [subsumption resolution 442,164] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
447. $false <- (4, 5, 45) [subsumption resolution 446,169] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
448. ~4 | ~5 | ~45 [avatar contradiction clause 447]
451. $false [avatar sat refutation 150,155,160,165,170,416,425,435,448] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:295,thDist:-295}
% SZS output end Proof for qc_lang2__t59_qc_lang2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 40
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 8
% 
% >>> Saturation
% Initial clauses: 39
% Generated clauses: 153
% Active clauses: 52
% Passive clauses: 72
% Final active clauses: 46
% Final passive clauses: 19
% Main loop iterations started: 53
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 14
% Bw demodulations: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 41
% Unit resulting resolution: 3
% Forward superposition: 1
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 55
% SAT solver unit clauses: 18
% SAT solver binary clauses: 23
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% backward subsumption index maintenance: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

