i: [1,53,0,2,1,1,0,0] l49_termord
a: [3,73,0,2,1,53]
s: 73
i: [1,51,0,2,1,1,0,0] l49_termord
a: [3,83,0,2,1,51]
s: 83
i: [1,50,0,2,1,1,0,0] l49_termord
a: [3,88,0,2,1,50]
s: 88
i: [1,49,0,2,1,1,0,0] l49_termord
a: [3,93,0,2,1,49]
s: 93
i: [1,54,0,3,1,1,0,0] l49_termord
a: [3,68,0,3,1,54]
s: 68
i: [1,48,0,3,1,1,0,0] l49_termord
a: [3,98,0,3,1,48]
s: 98
i: [1,47,0,4,1,1,0,0] l49_termord
a: [3,103,0,4,1,47]
s: 103
i: [1,52,0,6,1,1,0,0] l49_termord
a: [3,78,0,6,1,52]
s: 78
i: [1,62,0,3,1,0,0,2] reflexivity_r1_tarski
s: 62
i: [1,61,0,4,1,0,0,1] existence_m1_subset_1
s: 61
i: [1,59,0,7,2,0,0,2] t3_subset
s: 59
d: [2,105,1,3,1,1,40,59,103]
a: [3,111,1,3,1,105]
s: 111
i: [1,60,0,7,2,0,0,2] t3_subset
s: 60
d: [2,106,1,5,1,1,40,59,78]
a: [3,116,1,5,1,106]
s: 116
i: [1,55,0,8,3,0,0,1] t37_orders_1
s: 55
i: [1,63,0,8,2,0,0,2] cc1_relset_1
s: 63
d: [2,119,1,2,1,1,40,63,78]
a: [3,125,1,2,1,119]
s: 125
i: [1,56,0,11,4,0,0,1] t38_orders_1
s: 56
d: [2,128,2,8,3,1,40,56,111]
s: 128
i: [1,57,0,13,3,0,0,255] cc4_funct_2
s: 57
i: [1,64,0,13,3,0,0,255] cc1_funct_2
s: 64
d: [2,133,1,7,2,1,40,64,78]
d: [2,136,1,4,1,2,21,133,98]
a: [3,140,1,4,1,136]
s: 140
i: [1,58,0,19,6,0,0,255] t15_orders_1
s: 58
d: [2,142,1,13,5,1,40,58,78]
d: [2,145,1,10,4,2,21,142,98]
d: [2,146,1,8,3,3,21,145,83]
d: [2,147,1,6,2,4,21,146,88]
d: [2,148,1,4,1,5,21,147,93]
a: [3,152,1,4,1,148]
s: 152
d: [2,107,1,5,1,0,40,59,61]
s: 107
d: [2,121,1,6,1,0,40,63,61]
s: 121
d: [2,154,2,7,3,1,44,55,152]
d: [2,155,2,5,2,2,21,154,73]
d: [2,156,2,3,1,3,21,155,125]
a: [3,160,2,3,1,156]
s: 160
d: [2,163,3,0,0,4,52,125,68,160,128]
e: 163
d: [2,164,3,0,0,4,52,125,111,68,160,56]
e: 164
d: [2,165,3,5,2,2,40,160,128]
d: [2,168,3,2,1,3,21,165,68]
d: [2,169,3,0,0,4,21,168,125]
e: 169
f: 169,156,148,119,105,47,48,49,50,51,52,53,54
% Refutation found. Thanks to Tanya!
% SZS status Theorem for termord__l49_termord
% SZS output start Proof for termord__l49_termord
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) => (v3_orders_1(X2) => r3_orders_1(X2,X1)))) [input l49_termord]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) => (v3_orders_1(X2) => r3_orders_1(X2,X1)))) [negated conjecture 1]
4. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [input cc1_relset_1]
13. ! [X0] : ! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v8_relat_2(X1) & v4_relat_2(X1) & v1_relat_2(X1)) => k1_relat_1(X1) = X0) [input t15_orders_1]
14. ! [X0] : (v1_relat_1(X0) => (v3_orders_1(X0) => r3_orders_1(X0,k1_relat_1(X0)))) [input t37_orders_1]
15. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r3_orders_1(X0,X1)) => r3_orders_1(X0,X2))) [input t38_orders_1]
16. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
17. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) => (v3_orders_1(X2) => r3_orders_1(X2,X1)))) [flattening 2]
18. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : ((r1_tarski(X2,X1) & r3_orders_1(X0,X1)) => r3_orders_1(X0,X2))) [flattening 15]
20. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v8_relat_2(X1) & v4_relat_2(X1) & v1_relat_2(X1)) => k1_relat_1(X1) = X0) [flattening 13]
21. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 16]
23. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [flattening 4]
28. ? [X0,X1] : (? [X2] : ((~r3_orders_1(X2,X1) & v3_orders_1(X2)) & (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 17]
29. ? [X0,X1] : (? [X2] : (~r3_orders_1(X2,X1) & v3_orders_1(X2) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) & m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 28]
30. ! [X0] : ((r3_orders_1(X0,k1_relat_1(X0)) | ~v3_orders_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 14]
31. ! [X0] : (r3_orders_1(X0,k1_relat_1(X0)) | ~v3_orders_1(X0) | ~v1_relat_1(X0)) [flattening 30]
32. ! [X0] : (! [X1,X2] : (r3_orders_1(X0,X2) | (~r1_tarski(X2,X1) | ~r3_orders_1(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 18]
33. ! [X0] : (! [X1,X2] : (r3_orders_1(X0,X2) | ~r1_tarski(X2,X1) | ~r3_orders_1(X0,X1)) | ~v1_relat_1(X0)) [flattening 32]
36. ! [X0,X1] : (k1_relat_1(X1) = X0 | (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~v1_partfun1(X1,X0) | ~v8_relat_2(X1) | ~v4_relat_2(X1) | ~v1_relat_2(X1))) [ennf transformation 20]
37. ! [X0,X1] : (k1_relat_1(X1) = X0 | ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~v1_partfun1(X1,X0) | ~v8_relat_2(X1) | ~v4_relat_2(X1) | ~v1_relat_2(X1)) [flattening 36]
38. ! [X0,X1,X2] : (v1_relat_1(X2) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 23]
41. ? [X0,X1] : (? [X2] : (~r3_orders_1(X2,X1) & v3_orders_1(X2) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (? [X2] : (~r3_orders_1(X2,sK1) & v3_orders_1(X2) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,sK0)) & m1_subset_1(sK1,k1_zfmisc_1(sK0))) [choice axiom]
42. ? [X2] : (~r3_orders_1(X2,sK1) & v3_orders_1(X2) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,sK0)) => (~r3_orders_1(sK2,sK1) & v3_orders_1(sK2) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v8_relat_2(sK2) & v4_relat_2(sK2) & v1_relat_2(sK2) & v1_partfun1(sK2,sK0)) [choice axiom]
43. (~r3_orders_1(sK2,sK1) & v3_orders_1(sK2) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v8_relat_2(sK2) & v4_relat_2(sK2) & v1_relat_2(sK2) & v1_partfun1(sK2,sK0)) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) [skolemisation 29,42,41]
44. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 21]
47. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 43] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. v1_partfun1(sK2,sK0) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. v1_relat_2(sK2) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. v4_relat_2(sK2) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. v8_relat_2(sK2) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) [cnf transformation 43] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
53. v3_orders_1(sK2) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. ~r3_orders_1(sK2,sK1) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. r3_orders_1(X0,k1_relat_1(X0)) | ~v3_orders_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. ~r1_tarski(X2,X1) | r3_orders_1(X0,X2) | ~r3_orders_1(X0,X1) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | k1_relat_1(X1) = X0 | ~v1_partfun1(X1,X0) | ~v8_relat_2(X1) | ~v4_relat_2(X1) | ~v1_relat_2(X1) [cnf transformation 37] {a:0,w:19,wCS:7600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
59. ~m1_subset_1(X0,k1_zfmisc_1(X1)) | r1_tarski(X0,X1) [cnf transformation 44] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
63. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v1_relat_1(X2) [cnf transformation 38] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
66. 1 <=> r3_orders_1(sK2,sK1) [avatar definition]
68. ~r3_orders_1(sK2,sK1) <- (~1) [avatar component clause 66] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
69. ~1 [avatar split clause 54,66]
71. 2 <=> v3_orders_1(sK2) [avatar definition]
73. v3_orders_1(sK2) <- (2) [avatar component clause 71] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
74. 2 [avatar split clause 53,71]
76. 3 <=> m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) [avatar definition]
78. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) <- (3) [avatar component clause 76] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
79. 3 [avatar split clause 52,76]
81. 4 <=> v8_relat_2(sK2) [avatar definition]
83. v8_relat_2(sK2) <- (4) [avatar component clause 81] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
84. 4 [avatar split clause 51,81]
86. 5 <=> v4_relat_2(sK2) [avatar definition]
88. v4_relat_2(sK2) <- (5) [avatar component clause 86] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. 5 [avatar split clause 50,86]
91. 6 <=> v1_relat_2(sK2) [avatar definition]
93. v1_relat_2(sK2) <- (6) [avatar component clause 91] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
94. 6 [avatar split clause 49,91]
96. 7 <=> v1_partfun1(sK2,sK0) [avatar definition]
98. v1_partfun1(sK2,sK0) <- (7) [avatar component clause 96] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
99. 7 [avatar split clause 48,96]
101. 8 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
103. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (8) [avatar component clause 101] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. 8 [avatar split clause 47,101]
105. r1_tarski(sK1,sK0) <- (8) [resolution 59,103] {a:1,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
109. 9 <=> r1_tarski(sK1,sK0) [avatar definition]
111. r1_tarski(sK1,sK0) <- (9) [avatar component clause 109] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
112. 9 | ~8 [avatar split clause 105,101,109]
119. v1_relat_1(sK2) <- (3) [resolution 63,78] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
123. 11 <=> v1_relat_1(sK2) [avatar definition]
125. v1_relat_1(sK2) <- (11) [avatar component clause 123] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
126. 11 | ~3 [avatar split clause 119,76,123]
128. ~r3_orders_1(X2,sK0) | r3_orders_1(X2,sK1) | ~v1_relat_1(X2) <- (9) [resolution 56,111] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
142. sK0 = k1_relat_1(sK2) | ~v1_partfun1(sK2,sK0) | ~v8_relat_2(sK2) | ~v4_relat_2(sK2) | ~v1_relat_2(sK2) <- (3) [resolution 58,78] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
145. sK0 = k1_relat_1(sK2) | ~v8_relat_2(sK2) | ~v4_relat_2(sK2) | ~v1_relat_2(sK2) <- (3, 7) [subsumption resolution 142,98] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
146. sK0 = k1_relat_1(sK2) | ~v4_relat_2(sK2) | ~v1_relat_2(sK2) <- (3, 4, 7) [subsumption resolution 145,83] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
147. sK0 = k1_relat_1(sK2) | ~v1_relat_2(sK2) <- (3, 4, 5, 7) [subsumption resolution 146,88] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
148. sK0 = k1_relat_1(sK2) <- (3, 4, 5, 6, 7) [subsumption resolution 147,93] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
150. 13 <=> sK0 = k1_relat_1(sK2) [avatar definition]
152. sK0 = k1_relat_1(sK2) <- (13) [avatar component clause 150] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
153. 13 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 148,96,91,86,81,76,150]
154. r3_orders_1(sK2,sK0) | ~v3_orders_1(sK2) | ~v1_relat_1(sK2) <- (13) [superposition 55,152] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
155. r3_orders_1(sK2,sK0) | ~v1_relat_1(sK2) <- (2, 13) [subsumption resolution 154,73] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
156. r3_orders_1(sK2,sK0) <- (2, 11, 13) [subsumption resolution 155,125] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
158. 14 <=> r3_orders_1(sK2,sK0) [avatar definition]
160. r3_orders_1(sK2,sK0) <- (14) [avatar component clause 158] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
161. 14 | ~2 | ~11 | ~13 [avatar split clause 156,150,123,71,158]
165. r3_orders_1(sK2,sK1) | ~v1_relat_1(sK2) <- (9, 14) [resolution 160,128] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
168. ~v1_relat_1(sK2) <- (~1, 9, 14) [subsumption resolution 165,68] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
169. $false <- (~1, 9, 11, 14) [subsumption resolution 168,125] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
170. 1 | ~9 | ~11 | ~14 [avatar contradiction clause 169]
171. $false [avatar sat refutation 69,74,79,84,89,94,99,104,112,126,153,161,170] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:120,thDist:-120}
% SZS output end Proof for termord__l49_termord
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 18
% Generated clauses: 65
% Active clauses: 27
% Passive clauses: 36
% Final active clauses: 27
% Final passive clauses: 9
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 21
% Unit resulting resolution: 2
% Backward superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 19
% SAT solver unit clauses: 8
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

