i: [1,59,0,2,1,1,0,0] t35_rfunct_1
a: [3,97,0,2,1,59]
s: 97
i: [1,58,0,2,1,1,0,0] t35_rfunct_1
a: [3,102,0,2,1,58]
s: 102
i: [1,57,0,2,1,1,0,0] t35_rfunct_1
a: [3,107,0,2,1,57]
s: 107
i: [1,56,0,2,1,1,0,0] t35_rfunct_1
a: [3,112,0,2,1,56]
s: 112
i: [1,55,0,2,1,1,0,0] t35_rfunct_1
a: [3,117,0,2,1,55]
s: 117
i: [1,54,0,2,1,1,0,0] t35_rfunct_1
a: [3,122,0,2,1,54]
s: 122
i: [1,88,0,2,1,0,0,1] rc1_funct_1
a: [3,127,0,2,1,88]
s: 127
i: [1,87,0,2,1,0,0,1] rc1_funct_1
a: [3,132,0,2,1,87]
s: 132
i: [1,60,0,12,1,1,0,0] t35_rfunct_1
a: [3,92,0,12,1,60]
s: 92
i: [1,70,0,6,2,0,0,1] dt_k5_relat_1
s: 70
i: [1,61,0,8,2,0,0,1] rd4_relat_1
s: 61
i: [1,72,0,8,3,0,0,1] fc43_valued_0
s: 72
i: [1,73,0,8,3,0,0,1] fc43_valued_0
s: 73
i: [1,85,0,8,3,0,0,1] fc8_funct_1
s: 85
i: [1,86,0,8,3,0,0,1] fc8_funct_1
s: 86
i: [1,63,0,9,4,0,0,1] dt_k4_rfunct_1
s: 63
i: [1,64,0,9,4,0,0,1] dt_k4_rfunct_1
s: 64
i: [1,65,0,9,4,0,0,1] fc3_rfunct_1
s: 65
i: [1,66,0,9,4,0,0,1] fc3_rfunct_1
s: 66
i: [1,67,0,9,4,0,0,1] fc3_rfunct_1
s: 67
i: [1,71,0,11,2,0,0,1] rd5_relat_1
s: 71
i: [1,62,0,15,4,0,0,1] t26_rfunct_1
s: 62
i: [1,75,0,16,7,0,0,1] dt_k1_rfunct_1
s: 75
d: [2,146,1,16,4,1,44,92,62]
d: [2,155,1,14,3,2,21,146,97]
d: [2,156,1,12,2,3,21,155,102]
d: [2,157,1,10,1,4,21,156,107]
a: [3,161,1,10,1,157]
s: 161
i: [1,76,0,16,7,0,0,1] dt_k1_rfunct_1
s: 76
i: [1,77,0,16,7,0,0,1] fc1_rfunct_1
s: 77
d: [2,150,1,12,5,0,44,70,62]
d: [2,151,1,10,4,0,18,150]
s: 151
i: [1,78,0,16,7,0,0,1] fc1_rfunct_1
s: 78
d: [2,149,1,14,6,0,44,73,62]
d: [2,152,1,10,4,0,18,149]
s: 152
i: [1,79,0,16,7,0,0,1] fc1_rfunct_1
s: 79
i: [1,80,0,16,7,0,0,2] dt_k18_valued_1
s: 80
d: [2,148,1,14,6,0,44,86,62]
d: [2,153,1,10,4,0,18,148]
s: 153
i: [1,81,0,16,7,0,0,2] dt_k18_valued_1
s: 81
i: [1,82,0,16,7,0,0,2] fc33_valued_1
s: 82
i: [1,83,0,16,7,0,0,2] fc33_valued_1
s: 83
i: [1,84,0,16,7,0,0,2] fc33_valued_1
s: 84
d: [2,147,1,19,5,0,44,71,62]
d: [2,154,1,17,4,0,18,147]
s: 154
i: [1,74,0,19,7,0,0,2] commutativity_k18_valued_1
s: 74
i: [1,68,0,20,7,0,0,1] t31_rfunct_1
s: 68
i: [1,69,0,22,7,0,0,2] t27_rfunct_1
s: 69
d: [2,211,1,35,13,0,44,84,69]
d: [2,246,1,32,12,0,21,211,67]
d: [2,247,1,29,11,0,21,246,67]
d: [2,248,1,26,10,0,21,247,64]
d: [2,249,1,23,9,0,21,248,63]
d: [2,250,1,20,8,0,21,249,64]
d: [2,251,1,17,7,0,21,250,63]
s: 251
d: [2,212,1,35,13,0,44,81,69]
d: [2,252,1,32,12,0,21,212,64]
d: [2,253,1,29,11,0,21,252,64]
d: [2,254,1,26,10,0,21,253,67]
d: [2,255,1,23,9,0,21,254,63]
d: [2,256,1,20,8,0,21,255,67]
d: [2,257,1,17,7,0,21,256,63]
s: 257
d: [2,213,1,35,13,0,44,80,69]
d: [2,258,1,32,12,0,21,213,63]
d: [2,259,1,29,11,0,21,258,63]
d: [2,260,1,26,10,0,21,259,67]
d: [2,261,1,23,9,0,21,260,64]
d: [2,262,1,20,8,0,21,261,67]
d: [2,263,1,17,7,0,21,262,64]
s: 263
d: [2,266,2,32,13,0,44,251,68]
d: [2,268,2,26,10,0,18,266]
d: [2,271,2,23,9,0,21,268,67]
d: [2,272,2,20,8,0,21,271,64]
d: [2,273,2,17,7,0,21,272,63]
s: 273
d: [2,179,1,35,13,0,44,68,74]
d: [2,192,1,29,10,0,18,179]
d: [2,193,1,26,9,0,21,192,67]
d: [2,194,1,23,8,0,21,193,64]
d: [2,195,1,20,7,0,21,194,63]
s: 195
d: [2,180,1,35,13,0,44,68,74]
d: [2,191,1,29,10,0,18,180]
d: [2,196,1,26,9,0,21,191,67]
d: [2,197,1,23,8,0,21,196,64]
d: [2,198,1,20,7,0,21,197,63]
s: 198
d: [2,282,2,32,13,0,44,257,68]
d: [2,284,2,26,10,0,18,282]
d: [2,287,2,23,9,0,21,284,64]
d: [2,288,2,20,8,0,21,287,67]
d: [2,289,2,17,7,0,21,288,63]
s: 289
d: [2,181,1,35,13,0,44,74,68]
d: [2,190,1,29,10,0,18,181]
d: [2,199,1,26,9,0,21,190,67]
d: [2,200,1,23,8,0,21,199,64]
d: [2,201,1,20,7,0,21,200,63]
s: 201
d: [2,298,2,32,13,0,44,263,68]
d: [2,300,2,26,10,0,18,298]
d: [2,303,2,23,9,0,21,300,63]
d: [2,304,2,20,8,0,21,303,67]
d: [2,305,2,17,7,0,21,304,64]
s: 305
d: [2,182,1,35,13,0,44,74,68]
d: [2,189,1,29,10,0,18,182]
d: [2,202,1,26,9,0,21,189,67]
d: [2,203,1,23,8,0,21,202,64]
d: [2,204,1,20,7,0,21,203,63]
s: 204
d: [2,205,1,36,13,0,44,69,68]
d: [2,215,1,30,10,0,18,205]
d: [2,216,1,27,9,0,21,215,67]
d: [2,217,1,24,8,0,21,216,64]
d: [2,218,1,21,7,0,21,217,63]
s: 218
d: [2,362,2,25,7,1,44,161,218]
d: [2,375,2,23,6,2,21,362,112]
d: [2,376,2,21,5,3,21,375,117]
d: [2,377,2,19,4,4,21,376,122]
a: [3,381,2,3,1,377]
s: 381
d: [2,404,3,0,0,4,52,107,102,97,381,63]
e: 404
d: [2,405,3,6,3,1,40,381,63]
d: [2,407,3,4,2,2,21,405,97]
d: [2,408,3,2,1,3,21,407,102]
d: [2,409,3,0,0,4,21,408,107]
e: 409
r: 381
d: [2,208,1,36,13,0,44,68,69]
d: [2,214,1,30,10,0,18,208]
d: [2,231,1,27,9,0,21,214,67]
d: [2,232,1,24,8,0,21,231,64]
d: [2,233,1,21,7,0,21,232,63]
s: 233
a: [3,380,2,3,1,377]
s: 380
d: [2,206,1,40,13,0,44,69,74]
d: [2,219,1,37,12,0,21,206,67]
d: [2,220,1,34,11,0,21,219,64]
d: [2,221,1,31,10,0,21,220,63]
d: [2,222,1,28,9,0,21,221,67]
d: [2,223,1,25,8,0,21,222,64]
d: [2,224,1,22,7,0,21,223,63]
s: 224
d: [2,207,1,40,13,0,44,69,74]
d: [2,225,1,37,12,0,21,207,67]
d: [2,226,1,34,11,0,21,225,64]
d: [2,227,1,31,10,0,21,226,63]
d: [2,228,1,28,9,0,21,227,67]
d: [2,229,1,25,8,0,21,228,64]
d: [2,230,1,22,7,0,21,229,63]
s: 230
a: [3,385,2,3,1,377]
s: 385
d: [2,457,3,0,0,4,52,107,102,97,385,64]
e: 457
d: [2,458,3,6,3,1,40,385,64]
d: [2,460,3,4,2,2,21,458,97]
d: [2,461,3,2,1,3,21,460,102]
d: [2,462,3,0,0,4,21,461,107]
e: 462
r: 385
d: [2,209,1,40,13,0,44,74,69]
d: [2,234,1,37,12,0,21,209,67]
d: [2,235,1,34,11,0,21,234,64]
d: [2,236,1,31,10,0,21,235,63]
d: [2,237,1,28,9,0,21,236,67]
d: [2,238,1,25,8,0,21,237,64]
d: [2,239,1,22,7,0,21,238,63]
s: 239
a: [3,384,2,3,1,377]
s: 384
d: [2,210,1,40,13,0,44,74,69]
d: [2,240,1,37,12,0,21,210,67]
d: [2,241,1,34,11,0,21,240,64]
d: [2,242,1,31,10,0,21,241,63]
d: [2,243,1,28,9,0,21,242,67]
d: [2,244,1,25,8,0,21,243,64]
d: [2,245,1,22,7,0,21,244,63]
s: 245
a: [3,389,2,3,1,377]
s: 389
d: [2,464,3,0,0,4,52,107,102,97,389,67]
e: 464
d: [2,465,3,6,3,1,40,389,67]
d: [2,467,3,4,2,2,21,465,97]
d: [2,468,3,2,1,3,21,467,102]
d: [2,469,3,0,0,4,21,468,107]
e: 469
r: 389
a: [3,388,2,3,1,377]
s: 388
a: [3,393,2,10,1,377]
s: 393
d: [2,471,3,21,7,1,44,393,195]
d: [2,474,3,12,6,1,19,471]
d: [2,475,3,10,5,2,21,474,97]
d: [2,476,3,8,4,3,21,475,102]
d: [2,477,3,6,3,4,21,476,107]
d: [2,478,3,4,2,5,21,477,112]
d: [2,479,3,2,1,6,21,478,117]
d: [2,480,3,0,0,7,21,479,122]
e: 480
f: 480,469,462,409,377,157,54,55,56,57,58,59,60
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rfunct_1__t35_rfunct_1
% SZS output start Proof for rfunct_1__t35_rfunct_1
1. ! [X0] : ((v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => k4_rfunct_1(k1_rfunct_1(X0,X1)) = k1_rfunct_1(k5_relat_1(X1,k9_xtuple_0(k4_rfunct_1(X1))),X0))) [input t35_rfunct_1]
2. ~! [X0] : ((v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => k4_rfunct_1(k1_rfunct_1(X0,X1)) = k1_rfunct_1(k5_relat_1(X1,k9_xtuple_0(k4_rfunct_1(X1))),X0))) [negated conjecture 1]
3. ! [X0,X1] : ((v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1) & v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => k18_valued_1(X0,X1) = k18_valued_1(X1,X0)) [input commutativity_k18_valued_1]
6. ! [X0] : ((v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k4_rfunct_1(X0)) & v1_relat_1(k4_rfunct_1(X0)))) [input dt_k4_rfunct_1]
11. ! [X0] : ((v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v1_valued_0(k4_rfunct_1(X0)) & v1_funct_1(k4_rfunct_1(X0)) & v1_relat_1(k4_rfunct_1(X0)))) [input fc3_rfunct_1]
17. ! [X0] : ((v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => k4_rfunct_1(k4_rfunct_1(X0)) = k5_relat_1(X0,k9_xtuple_0(k4_rfunct_1(X0)))) [input t26_rfunct_1]
18. ! [X0] : ((v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => k4_rfunct_1(k18_valued_1(X0,X1)) = k18_valued_1(k4_rfunct_1(X0),k4_rfunct_1(X1)))) [input t27_rfunct_1]
19. ! [X0] : ((v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => k1_rfunct_1(X0,X1) = k18_valued_1(X0,k4_rfunct_1(X1)))) [input t31_rfunct_1]
20. ? [X0] : (? [X1] : (k4_rfunct_1(k1_rfunct_1(X0,X1)) != k1_rfunct_1(k5_relat_1(X1,k9_xtuple_0(k4_rfunct_1(X1))),X0) & (v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1))) & (v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
21. ? [X0] : (? [X1] : (k4_rfunct_1(k1_rfunct_1(X0,X1)) != k1_rfunct_1(k5_relat_1(X1,k9_xtuple_0(k4_rfunct_1(X1))),X0) & v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 20]
23. ! [X0] : (k4_rfunct_1(k4_rfunct_1(X0)) = k5_relat_1(X0,k9_xtuple_0(k4_rfunct_1(X0))) | (~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 17]
24. ! [X0] : (k4_rfunct_1(k4_rfunct_1(X0)) = k5_relat_1(X0,k9_xtuple_0(k4_rfunct_1(X0))) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 23]
25. ! [X0] : ((v1_funct_1(k4_rfunct_1(X0)) & v1_relat_1(k4_rfunct_1(X0))) | (~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 6]
26. ! [X0] : ((v1_funct_1(k4_rfunct_1(X0)) & v1_relat_1(k4_rfunct_1(X0))) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 25]
27. ! [X0] : ((v1_valued_0(k4_rfunct_1(X0)) & v1_funct_1(k4_rfunct_1(X0)) & v1_relat_1(k4_rfunct_1(X0))) | (~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 11]
28. ! [X0] : ((v1_valued_0(k4_rfunct_1(X0)) & v1_funct_1(k4_rfunct_1(X0)) & v1_relat_1(k4_rfunct_1(X0))) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 27]
29. ! [X0] : (! [X1] : (k1_rfunct_1(X0,X1) = k18_valued_1(X0,k4_rfunct_1(X1)) | (~v1_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 19]
30. ! [X0] : (! [X1] : (k1_rfunct_1(X0,X1) = k18_valued_1(X0,k4_rfunct_1(X1)) | ~v1_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 29]
31. ! [X0] : (! [X1] : (k4_rfunct_1(k18_valued_1(X0,X1)) = k18_valued_1(k4_rfunct_1(X0),k4_rfunct_1(X1)) | (~v1_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 18]
32. ! [X0] : (! [X1] : (k4_rfunct_1(k18_valued_1(X0,X1)) = k18_valued_1(k4_rfunct_1(X0),k4_rfunct_1(X1)) | ~v1_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 31]
37. ! [X0,X1] : (k18_valued_1(X0,X1) = k18_valued_1(X1,X0) | (~v1_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 3]
38. ! [X0,X1] : (k18_valued_1(X0,X1) = k18_valued_1(X1,X0) | ~v1_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 37]
49. ? [X0] : (? [X1] : (k4_rfunct_1(k1_rfunct_1(X0,X1)) != k1_rfunct_1(k5_relat_1(X1,k9_xtuple_0(k4_rfunct_1(X1))),X0) & v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (k4_rfunct_1(k1_rfunct_1(sK0,X1)) != k1_rfunct_1(k5_relat_1(X1,k9_xtuple_0(k4_rfunct_1(X1))),sK0) & v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_valued_0(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
50. ? [X1] : (k4_rfunct_1(k1_rfunct_1(sK0,X1)) != k1_rfunct_1(k5_relat_1(X1,k9_xtuple_0(k4_rfunct_1(X1))),sK0) & v1_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => (k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k1_rfunct_1(k5_relat_1(sK1,k9_xtuple_0(k4_rfunct_1(sK1))),sK0) & v1_valued_0(sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
51. (k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k1_rfunct_1(k5_relat_1(sK1,k9_xtuple_0(k4_rfunct_1(sK1))),sK0) & v1_valued_0(sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_valued_0(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 21,50,49]
54. v1_relat_1(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. v1_funct_1(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. v1_valued_0(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
57. v1_relat_1(sK1) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. v1_funct_1(sK1) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
59. v1_valued_0(sK1) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
60. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k1_rfunct_1(k5_relat_1(sK1,k9_xtuple_0(k4_rfunct_1(sK1))),sK0) [cnf transformation 51] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
62. k4_rfunct_1(k4_rfunct_1(X0)) = k5_relat_1(X0,k9_xtuple_0(k4_rfunct_1(X0))) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. v1_relat_1(k4_rfunct_1(X0)) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. v1_funct_1(k4_rfunct_1(X0)) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. v1_valued_0(k4_rfunct_1(X0)) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. k1_rfunct_1(X0,X1) = k18_valued_1(X0,k4_rfunct_1(X1)) | ~v1_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. k4_rfunct_1(k18_valued_1(X0,X1)) = k18_valued_1(k4_rfunct_1(X0),k4_rfunct_1(X1)) | ~v1_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 32] {a:0,w:22,wCS:8800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
74. k18_valued_1(X0,X1) = k18_valued_1(X1,X0) | ~v1_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 38] {a:0,w:19,wCS:7600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
90. 1 <=> k4_rfunct_1(k1_rfunct_1(sK0,sK1)) = k1_rfunct_1(k5_relat_1(sK1,k9_xtuple_0(k4_rfunct_1(sK1))),sK0) [avatar definition]
92. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k1_rfunct_1(k5_relat_1(sK1,k9_xtuple_0(k4_rfunct_1(sK1))),sK0) <- (~1) [avatar component clause 90] {a:0,w:12,wCS:1200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
93. ~1 [avatar split clause 60,90]
95. 2 <=> v1_valued_0(sK1) [avatar definition]
97. v1_valued_0(sK1) <- (2) [avatar component clause 95] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. 2 [avatar split clause 59,95]
100. 3 <=> v1_funct_1(sK1) [avatar definition]
102. v1_funct_1(sK1) <- (3) [avatar component clause 100] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
103. 3 [avatar split clause 58,100]
105. 4 <=> v1_relat_1(sK1) [avatar definition]
107. v1_relat_1(sK1) <- (4) [avatar component clause 105] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
108. 4 [avatar split clause 57,105]
110. 5 <=> v1_valued_0(sK0) [avatar definition]
112. v1_valued_0(sK0) <- (5) [avatar component clause 110] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. 5 [avatar split clause 56,110]
115. 6 <=> v1_funct_1(sK0) [avatar definition]
117. v1_funct_1(sK0) <- (6) [avatar component clause 115] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. 6 [avatar split clause 55,115]
120. 7 <=> v1_relat_1(sK0) [avatar definition]
122. v1_relat_1(sK0) <- (7) [avatar component clause 120] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
123. 7 [avatar split clause 54,120]
146. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k1_rfunct_1(k4_rfunct_1(k4_rfunct_1(sK1)),sK0) | ~v1_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1) [superposition 92,62] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
155. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k1_rfunct_1(k4_rfunct_1(k4_rfunct_1(sK1)),sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~1, 2) [subsumption resolution 146,97] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
156. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k1_rfunct_1(k4_rfunct_1(k4_rfunct_1(sK1)),sK0) | ~v1_relat_1(sK1) <- (~1, 2, 3) [subsumption resolution 155,102] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
157. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k1_rfunct_1(k4_rfunct_1(k4_rfunct_1(sK1)),sK0) <- (~1, 2, 3, 4) [subsumption resolution 156,107] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
159. 10 <=> k4_rfunct_1(k1_rfunct_1(sK0,sK1)) = k1_rfunct_1(k4_rfunct_1(k4_rfunct_1(sK1)),sK0) [avatar definition]
161. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k1_rfunct_1(k4_rfunct_1(k4_rfunct_1(sK1)),sK0) <- (~10) [avatar component clause 159] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
162. ~10 | 1 | ~2 | ~3 | ~4 [avatar split clause 157,105,100,95,90,159]
179. k18_valued_1(k4_rfunct_1(X3),X2) = k1_rfunct_1(X2,X3) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_valued_0(k4_rfunct_1(X3)) | ~v1_funct_1(k4_rfunct_1(X3)) | ~v1_relat_1(k4_rfunct_1(X3)) [superposition 68,74] {a:1,w:35,wCS:14000,sine:1,thAx:0,allAx:2,thDist:-2}
192. k18_valued_1(k4_rfunct_1(X3),X2) = k1_rfunct_1(X2,X3) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_valued_0(k4_rfunct_1(X3)) | ~v1_funct_1(k4_rfunct_1(X3)) | ~v1_relat_1(k4_rfunct_1(X3)) [duplicate literal removal 179] {a:1,w:29,wCS:11600,sine:1,thAx:0,allAx:2,thDist:-2}
193. k18_valued_1(k4_rfunct_1(X3),X2) = k1_rfunct_1(X2,X3) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(k4_rfunct_1(X3)) | ~v1_relat_1(k4_rfunct_1(X3)) [subsumption resolution 192,67] {a:1,w:26,wCS:10400,sine:1,thAx:0,allAx:2,thDist:-2}
194. k18_valued_1(k4_rfunct_1(X3),X2) = k1_rfunct_1(X2,X3) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(k4_rfunct_1(X3)) [subsumption resolution 193,64] {a:1,w:23,wCS:9200,sine:1,thAx:0,allAx:2,thDist:-2}
195. k18_valued_1(k4_rfunct_1(X3),X2) = k1_rfunct_1(X2,X3) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [subsumption resolution 194,63] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
205. k1_rfunct_1(k4_rfunct_1(X2),X3) = k4_rfunct_1(k18_valued_1(X2,X3)) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(k4_rfunct_1(X2)) | ~v1_funct_1(k4_rfunct_1(X2)) | ~v1_relat_1(k4_rfunct_1(X2)) [superposition 69,68] {a:1,w:36,wCS:14400,sine:1,thAx:0,allAx:2,thDist:-2}
215. k1_rfunct_1(k4_rfunct_1(X2),X3) = k4_rfunct_1(k18_valued_1(X2,X3)) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_valued_0(k4_rfunct_1(X2)) | ~v1_funct_1(k4_rfunct_1(X2)) | ~v1_relat_1(k4_rfunct_1(X2)) [duplicate literal removal 205] {a:1,w:30,wCS:12000,sine:1,thAx:0,allAx:2,thDist:-2}
216. k1_rfunct_1(k4_rfunct_1(X2),X3) = k4_rfunct_1(k18_valued_1(X2,X3)) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(k4_rfunct_1(X2)) | ~v1_relat_1(k4_rfunct_1(X2)) [subsumption resolution 215,67] {a:1,w:27,wCS:10800,sine:1,thAx:0,allAx:2,thDist:-2}
217. k1_rfunct_1(k4_rfunct_1(X2),X3) = k4_rfunct_1(k18_valued_1(X2,X3)) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(k4_rfunct_1(X2)) [subsumption resolution 216,64] {a:1,w:24,wCS:9600,sine:1,thAx:0,allAx:2,thDist:-2}
218. k1_rfunct_1(k4_rfunct_1(X2),X3) = k4_rfunct_1(k18_valued_1(X2,X3)) | ~v1_valued_0(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [subsumption resolution 217,63] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
362. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k4_rfunct_1(k18_valued_1(k4_rfunct_1(sK1),sK0)) | ~v1_valued_0(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_valued_0(k4_rfunct_1(sK1)) | ~v1_funct_1(k4_rfunct_1(sK1)) | ~v1_relat_1(k4_rfunct_1(sK1)) <- (~10) [superposition 161,218] {a:2,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
375. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k4_rfunct_1(k18_valued_1(k4_rfunct_1(sK1),sK0)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_valued_0(k4_rfunct_1(sK1)) | ~v1_funct_1(k4_rfunct_1(sK1)) | ~v1_relat_1(k4_rfunct_1(sK1)) <- (5, ~10) [subsumption resolution 362,112] {a:2,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
376. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k4_rfunct_1(k18_valued_1(k4_rfunct_1(sK1),sK0)) | ~v1_relat_1(sK0) | ~v1_valued_0(k4_rfunct_1(sK1)) | ~v1_funct_1(k4_rfunct_1(sK1)) | ~v1_relat_1(k4_rfunct_1(sK1)) <- (5, 6, ~10) [subsumption resolution 375,117] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
377. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k4_rfunct_1(k18_valued_1(k4_rfunct_1(sK1),sK0)) | ~v1_valued_0(k4_rfunct_1(sK1)) | ~v1_funct_1(k4_rfunct_1(sK1)) | ~v1_relat_1(k4_rfunct_1(sK1)) <- (5, 6, 7, ~10) [subsumption resolution 376,122] {a:2,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
379. 11 <=> v1_relat_1(k4_rfunct_1(sK1)) [avatar definition]
381. ~v1_relat_1(k4_rfunct_1(sK1)) <- (~11) [avatar component clause 379] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
383. 12 <=> v1_funct_1(k4_rfunct_1(sK1)) [avatar definition]
385. ~v1_funct_1(k4_rfunct_1(sK1)) <- (~12) [avatar component clause 383] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
387. 13 <=> v1_valued_0(k4_rfunct_1(sK1)) [avatar definition]
389. ~v1_valued_0(k4_rfunct_1(sK1)) <- (~13) [avatar component clause 387] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
391. 14 <=> k4_rfunct_1(k1_rfunct_1(sK0,sK1)) = k4_rfunct_1(k18_valued_1(k4_rfunct_1(sK1),sK0)) [avatar definition]
393. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k4_rfunct_1(k18_valued_1(k4_rfunct_1(sK1),sK0)) <- (~14) [avatar component clause 391] {a:2,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
394. ~11 | ~12 | ~13 | ~14 | ~5 | ~6 | ~7 | 10 [avatar split clause 377,159,120,115,110,391,387,383,379]
405. ~v1_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~11) [resolution 381,63] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
407. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, ~11) [subsumption resolution 405,97] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
408. ~v1_relat_1(sK1) <- (2, 3, ~11) [subsumption resolution 407,102] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
409. $false <- (2, 3, 4, ~11) [subsumption resolution 408,107] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
410. ~2 | ~3 | ~4 | 11 [avatar contradiction clause 409]
458. ~v1_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~12) [resolution 385,64] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
460. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, ~12) [subsumption resolution 458,97] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
461. ~v1_relat_1(sK1) <- (2, 3, ~12) [subsumption resolution 460,102] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
462. $false <- (2, 3, 4, ~12) [subsumption resolution 461,107] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
463. ~2 | ~3 | ~4 | 12 [avatar contradiction clause 462]
465. ~v1_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~13) [resolution 389,67] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
467. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, ~13) [subsumption resolution 465,97] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
468. ~v1_relat_1(sK1) <- (2, 3, ~13) [subsumption resolution 467,102] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
469. $false <- (2, 3, 4, ~13) [subsumption resolution 468,107] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
470. ~2 | ~3 | ~4 | 13 [avatar contradiction clause 469]
471. k4_rfunct_1(k1_rfunct_1(sK0,sK1)) != k4_rfunct_1(k1_rfunct_1(sK0,sK1)) | ~v1_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_valued_0(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~14) [superposition 393,195] {a:3,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
474. ~v1_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_valued_0(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~14) [trivial inequality removal 471] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
475. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_valued_0(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, ~14) [subsumption resolution 474,97] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
476. ~v1_relat_1(sK1) | ~v1_valued_0(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3, ~14) [subsumption resolution 475,102] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
477. ~v1_valued_0(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3, 4, ~14) [subsumption resolution 476,107] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
478. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3, 4, 5, ~14) [subsumption resolution 477,112] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
479. ~v1_relat_1(sK0) <- (2, 3, 4, 5, 6, ~14) [subsumption resolution 478,117] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
480. $false <- (2, 3, 4, 5, 6, 7, ~14) [subsumption resolution 479,122] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
481. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | 14 [avatar contradiction clause 480]
494. $false [avatar sat refutation 93,98,103,108,113,118,123,162,394,410,463,470,481] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:134,thDist:-134}
% SZS output end Proof for rfunct_1__t35_rfunct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 35
% Generated clauses: 383
% Active clauses: 48
% Passive clauses: 82
% Final active clauses: 45
% Final passive clauses: 19
% Main loop iterations started: 63
% 
% >>> Simplifying Inferences
% Duplicate literals: 196
% Trivial inequalities: 1
% Fw subsumption resolutions: 177
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 53
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 3
% Forward superposition: 32
% Backward superposition: 58
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 4
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 9
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

