i: [1,38,0,2,1,1,0,0] t89_card_3
a: [3,72,0,2,1,38]
s: 72
i: [1,37,0,2,1,1,0,0] t89_card_3
a: [3,77,0,2,1,37]
s: 77
i: [1,40,0,3,1,1,0,0] t89_card_3
a: [3,62,0,3,1,40]
s: 62
i: [1,39,0,5,1,1,0,0] t89_card_3
a: [3,67,0,5,1,39]
s: 67
i: [1,56,0,2,1,0,0,255] rc1_ordinal1
a: [3,82,0,2,1,56]
s: 82
i: [1,57,0,2,1,0,0,1] rc1_card_1
a: [3,87,0,2,1,57]
s: 87
i: [1,41,0,3,1,0,0,1] dt_k2_card_1
s: 41
i: [1,42,0,3,1,0,0,2] dt_k1_card_1
s: 42
i: [1,52,0,3,1,0,0,2] reflexivity_r2_wellord2
s: 52
i: [1,44,0,4,2,0,0,255] cc1_ordinal1
s: 44
i: [1,45,0,4,2,0,0,255] cc1_ordinal1
s: 45
d: [2,90,1,2,1,1,40,44,82]
a: [3,94,1,2,1,90]
s: 94
i: [1,46,0,4,2,0,0,255] cc1_card_1
s: 46
d: [2,104,1,2,1,1,40,46,77]
a: [3,110,1,2,1,104]
s: 110
i: [1,43,0,6,1,0,0,1] projectivity_k1_card_1
s: 43
i: [1,47,0,6,3,0,0,255] cc2_ordinal1
s: 47
d: [2,105,1,2,1,1,40,46,72]
a: [3,115,1,2,1,105]
s: 115
i: [1,53,0,6,2,0,0,2] symmetry_r2_wellord2
s: 53
d: [2,122,2,2,1,1,40,110,45]
a: [3,127,2,2,1,122]
s: 127
i: [1,54,0,6,2,0,0,255] redefinition_r2_wellord2
s: 54
i: [1,55,0,6,2,0,0,255] redefinition_r2_wellord2
s: 55
d: [2,123,2,2,1,1,40,110,44]
a: [3,132,2,2,1,123]
s: 132
i: [1,58,0,7,2,0,0,1] d2_card_1
s: 58
d: [2,136,2,2,1,1,40,115,45]
a: [3,141,2,2,1,136]
s: 141
d: [2,151,0,4,1,0,21,58,42]
s: 151
i: [1,49,0,9,3,0,0,1] d2_card_1
s: 49
d: [2,137,2,2,1,1,40,115,44]
a: [3,146,2,2,1,137]
s: 146
i: [1,50,0,10,2,0,0,1] t88_card_3
s: 50
d: [2,96,1,2,1,1,40,45,82]
a: [3,100,1,2,1,96]
s: 100
i: [1,51,0,10,2,0,0,1] t88_card_3
s: 51
d: [2,106,1,2,1,1,40,46,87]
a: [3,120,1,2,1,106]
s: 120
d: [2,171,2,2,1,1,40,120,45]
a: [3,176,2,2,1,171]
s: 176
d: [2,168,1,10,2,1,44,51,67]
s: 168
d: [2,185,2,5,1,1,47,168]
a: [3,189,2,5,1,185]
s: 189
d: [2,169,1,10,2,1,44,51,67]
s: 169
d: [2,102,1,3,1,0,40,46,41]
s: 102
d: [2,195,3,4,1,1,44,151,189]
a: [3,203,3,4,1,195]
s: 203
d: [2,103,1,3,1,0,40,46,42]
s: 103
d: [2,208,4,4,1,1,40,203,54]
a: [3,214,4,4,1,208]
s: 214
d: [2,150,1,3,1,0,40,54,52]
s: 150
d: [2,153,1,4,1,0,40,151,54]
s: 153
d: [2,209,4,4,1,1,40,203,53]
a: [3,219,4,4,1,209]
s: 219
d: [2,154,1,4,1,0,40,151,53]
s: 154
d: [2,233,5,4,1,1,40,219,54]
a: [3,239,5,4,1,233]
s: 239
d: [2,159,1,9,3,0,44,43,49]
s: 159
d: [2,156,1,10,3,0,44,49,43]
s: 156
d: [2,259,6,8,3,1,44,239,159]
d: [2,273,6,6,2,2,21,259,77]
a: [3,276,6,3,1,273]
s: 276
d: [2,316,7,0,0,1,40,276,154]
e: 316
d: [2,315,7,0,0,1,40,276,52]
e: 315
d: [2,312,7,0,0,1,52,52,276]
e: 312
d: [2,313,7,0,0,1,52,150,276,55]
e: 313
d: [2,314,7,0,0,1,52,52,276,53]
e: 314
r: 276
d: [2,161,1,10,3,0,44,43,49]
s: 161
d: [2,262,4,8,3,1,44,203,159]
d: [2,294,4,6,2,2,21,262,77]
a: [3,298,4,3,1,294]
s: 298
d: [2,162,1,11,2,0,44,50,43]
s: 162
d: [2,164,1,11,2,0,44,50,43]
s: 164
d: [2,260,5,8,3,1,44,219,159]
d: [2,282,5,6,2,2,21,260,77]
a: [3,286,5,3,1,282]
s: 286
d: [2,157,1,13,5,0,44,49,49]
s: 157
d: [2,339,5,0,0,4,52,77,72,298,62,52,157]
e: 339
d: [2,340,6,0,0,4,52,72,77,286,62,52,157]
e: 340
d: [2,341,5,0,0,4,52,77,72,62,154,219,157]
e: 341
d: [2,342,5,0,0,4,52,72,77,62,219,154,157]
e: 342
d: [2,343,6,0,0,4,52,77,72,52,62,286,157]
e: 343
d: [2,344,5,0,0,4,52,72,77,52,62,298,157]
e: 344
f: 342,209,195,185,37,38,39,40
% Refutation found. Thanks to Tanya!
% SZS status Theorem for card_3__t89_card_3
% SZS output start Proof for card_3__t89_card_3
1. ! [X0] : (v1_card_1(X0) => ! [X1] : (v1_card_1(X1) => (k2_card_1(X0) = k2_card_1(X1) => X0 = X1))) [input t89_card_3]
2. ~! [X0] : (v1_card_1(X0) => ! [X1] : (v1_card_1(X1) => (k2_card_1(X0) = k2_card_1(X1) => X0 = X1))) [negated conjecture 1]
6. ! [X0] : ! [X1] : (v1_card_1(X1) => (k1_card_1(X0) = X1 <=> r2_wellord2(X0,X1))) [input d2_card_1]
7. ! [X0] : v1_card_1(k1_card_1(X0)) [input dt_k1_card_1]
14. ! [X0,X1] : (r2_wellord2(X0,X1) => r2_wellord2(X1,X0)) [input symmetry_r2_wellord2]
15. ! [X0] : ! [X1] : (k1_card_1(X0) = k1_card_1(X1) <=> k2_card_1(X0) = k2_card_1(X1)) [input t88_card_3]
16. ! [X0,X1] : (v1_card_1(X1) => (k1_card_1(X0) = X1 <=> r2_wellord2(X0,X1))) [flattening 6]
17. ! [X0,X1] : (k1_card_1(X0) = k1_card_1(X1) <=> k2_card_1(X0) = k2_card_1(X1)) [flattening 15]
19. ? [X0] : (? [X1] : ((X0 != X1 & k2_card_1(X0) = k2_card_1(X1)) & v1_card_1(X1)) & v1_card_1(X0)) [ennf transformation 2]
20. ? [X0] : (? [X1] : (X0 != X1 & k2_card_1(X0) = k2_card_1(X1) & v1_card_1(X1)) & v1_card_1(X0)) [flattening 19]
25. ! [X0,X1] : ((k1_card_1(X0) = X1 <=> r2_wellord2(X0,X1)) | ~v1_card_1(X1)) [ennf transformation 16]
26. ! [X0,X1] : (r2_wellord2(X1,X0) | ~r2_wellord2(X0,X1)) [ennf transformation 14]
27. ? [X0] : (? [X1] : (X0 != X1 & k2_card_1(X0) = k2_card_1(X1) & v1_card_1(X1)) & v1_card_1(X0)) => (? [X1] : (sK0 != X1 & k2_card_1(X1) = k2_card_1(sK0) & v1_card_1(X1)) & v1_card_1(sK0)) [choice axiom]
28. ? [X1] : (sK0 != X1 & k2_card_1(X1) = k2_card_1(sK0) & v1_card_1(X1)) => (sK0 != sK1 & k2_card_1(sK0) = k2_card_1(sK1) & v1_card_1(sK1)) [choice axiom]
29. (sK0 != sK1 & k2_card_1(sK0) = k2_card_1(sK1) & v1_card_1(sK1)) & v1_card_1(sK0) [skolemisation 20,28,27]
30. ! [X0,X1] : (((k1_card_1(X0) = X1 | ~r2_wellord2(X0,X1)) & (r2_wellord2(X0,X1) | k1_card_1(X0) != X1)) | ~v1_card_1(X1)) [nnf transformation 25]
31. ! [X0,X1] : ((k1_card_1(X0) = k1_card_1(X1) | k2_card_1(X0) != k2_card_1(X1)) & (k2_card_1(X0) = k2_card_1(X1) | k1_card_1(X0) != k1_card_1(X1))) [nnf transformation 17]
37. v1_card_1(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. v1_card_1(sK1) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. k2_card_1(sK0) = k2_card_1(sK1) [cnf transformation 29] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. sK0 != sK1 [cnf transformation 29] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. v1_card_1(k1_card_1(X0)) [cnf transformation 7] {a:0,w:3,wCS:1200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
48. r2_wellord2(X0,X1) | k1_card_1(X0) != X1 | ~v1_card_1(X1) [cnf transformation 30] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:1,thDist:-1}
49. k1_card_1(X0) = X1 | ~r2_wellord2(X0,X1) | ~v1_card_1(X1) [cnf transformation 30] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k2_card_1(X0) != k2_card_1(X1) | k1_card_1(X0) = k1_card_1(X1) [cnf transformation 31] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. ~r2_wellord2(X0,X1) | r2_wellord2(X1,X0) [cnf transformation 26] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
58. r2_wellord2(X0,k1_card_1(X0)) | ~v1_card_1(k1_card_1(X0)) [equality resolution 48] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
60. 1 <=> sK0 = sK1 [avatar definition]
62. sK0 != sK1 <- (~1) [avatar component clause 60] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
63. ~1 [avatar split clause 40,60]
65. 2 <=> k2_card_1(sK0) = k2_card_1(sK1) [avatar definition]
67. k2_card_1(sK0) = k2_card_1(sK1) <- (2) [avatar component clause 65] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
68. 2 [avatar split clause 39,65]
70. 3 <=> v1_card_1(sK1) [avatar definition]
72. v1_card_1(sK1) <- (3) [avatar component clause 70] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
73. 3 [avatar split clause 38,70]
75. 4 <=> v1_card_1(sK0) [avatar definition]
77. v1_card_1(sK0) <- (4) [avatar component clause 75] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
78. 4 [avatar split clause 37,75]
151. r2_wellord2(X0,k1_card_1(X0)) [subsumption resolution 58,42] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
154. r2_wellord2(k1_card_1(X1),X1) [resolution 151,53] {a:1,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
157. ~r2_wellord2(X2,X4) | X3 = X4 | ~v1_card_1(X4) | ~r2_wellord2(X2,X3) | ~v1_card_1(X3) [superposition 49,49] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
168. k2_card_1(X0) != k2_card_1(sK0) | k1_card_1(X0) = k1_card_1(sK1) <- (2) [superposition 51,67] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
185. k1_card_1(sK0) = k1_card_1(sK1) <- (2) [equality resolution 168] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
187. 18 <=> k1_card_1(sK0) = k1_card_1(sK1) [avatar definition]
189. k1_card_1(sK0) = k1_card_1(sK1) <- (18) [avatar component clause 187] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
190. 18 | ~2 [avatar split clause 185,65,187]
195. r2_wellord2(sK1,k1_card_1(sK0)) <- (18) [superposition 151,189] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
201. 19 <=> r2_wellord2(sK1,k1_card_1(sK0)) [avatar definition]
203. r2_wellord2(sK1,k1_card_1(sK0)) <- (19) [avatar component clause 201] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
204. 19 | ~18 [avatar split clause 195,187,201]
209. r2_wellord2(k1_card_1(sK0),sK1) <- (19) [resolution 203,53] {a:4,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
217. 21 <=> r2_wellord2(k1_card_1(sK0),sK1) [avatar definition]
219. r2_wellord2(k1_card_1(sK0),sK1) <- (21) [avatar component clause 217] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
220. 21 | ~19 [avatar split clause 209,201,217]
342. $false <- (~1, 3, 4, 21) [unit resulting resolution 72,77,62,219,154,157] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
355. 1 | ~3 | ~4 | ~21 [avatar contradiction clause 342]
362. $false [avatar sat refutation 63,68,73,78,190,204,220,355] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:407,thDist:-407}
% SZS output end Proof for card_3__t89_card_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 194
% Active clauses: 49
% Passive clauses: 83
% Final active clauses: 48
% Final passive clauses: 30
% Main loop iterations started: 53
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 12
% Fw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 3
% Forward subsumptions: 45
% 
% >>> Generating Inferences
% Binary resolution: 52
% Unit resulting resolution: 9
% Forward superposition: 38
% Backward superposition: 25
% Self superposition: 2
% Equality resolution: 4
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% Sat splitting refutations: 11
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 45
% SAT solver unit clauses: 11
% SAT solver binary clauses: 19
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

