i: [1,32,0,14,5,1,0,0] t20_gate_1
a: [3,50,0,2,1,32]
s: 50
a: [3,45,0,6,1,32]
s: 45
i: [1,33,0,2,1,0,0,255] fc1_xboole_0
a: [3,71,0,2,1,33]
s: 71
i: [1,41,0,2,1,0,0,1] rc2_xboole_0
a: [3,76,0,2,1,41]
s: 76
i: [1,42,0,2,1,0,0,1] rc1_xboole_0
a: [3,81,0,2,1,42]
s: 81
i: [1,34,0,5,2,0,0,255] fc2_gate_1
s: 34
i: [1,35,0,5,2,0,0,255] fc1_gate_1
s: 35
i: [1,37,0,9,2,0,0,255] d14_gate_1
s: 37
d: [2,84,1,4,2,1,44,45,37]
d: [2,85,1,2,1,2,21,84,71]
d: [2,86,1,0,0,3,21,85,50]
e: 86
r: 50
a: [3,49,0,2,1,32]
s: 49
a: [3,54,0,2,1,32]
s: 54
i: [1,38,0,9,2,0,0,255] d14_gate_1
s: 38
d: [2,89,1,4,2,1,44,45,38]
d: [2,90,1,2,1,2,21,89,71]
d: [2,91,1,0,0,3,21,90,54]
e: 91
r: 54
a: [3,53,0,2,1,32]
s: 53
a: [3,58,0,2,1,32]
s: 58
i: [1,39,0,9,2,0,0,255] d14_gate_1
s: 39
d: [2,94,1,4,2,1,44,45,39]
d: [2,95,1,2,1,2,21,94,71]
d: [2,96,1,0,0,3,21,95,58]
e: 96
r: 58
a: [3,57,0,2,1,32]
s: 57
a: [3,62,0,2,1,32]
s: 62
i: [1,40,0,9,2,0,0,255] d14_gate_1
s: 40
d: [2,99,1,4,2,1,44,45,40]
d: [2,100,1,2,1,2,21,99,71]
d: [2,101,1,0,0,3,21,100,62]
e: 101
r: 45
r: 62
a: [3,61,0,2,1,32]
s: 61
a: [3,46,0,6,1,32]
s: 46
i: [1,36,0,16,5,0,0,255] d14_gate_1
s: 36
d: [2,115,1,11,5,1,44,46,36]
d: [2,116,1,9,4,2,21,115,61]
d: [2,117,1,7,3,3,21,116,57]
d: [2,118,1,5,2,4,21,117,53]
d: [2,119,1,3,1,5,21,118,49]
a: [3,123,1,3,1,119]
s: 123
d: [2,125,2,0,0,2,52,71,123,35]
e: 125
d: [2,126,2,2,1,1,40,123,35]
d: [2,128,2,0,0,2,21,126,71]
e: 128
i: [1,31,0,8,2,1,0,0] t20_gate_1
i: [1,30,0,8,2,1,0,0] t20_gate_1
i: [1,29,0,8,2,1,0,0] t20_gate_1
i: [1,28,0,8,2,1,0,0] t20_gate_1
f: 128,119,101,95,90,85,33,28,29,30,31,32
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t20_gate_1
% SZS output start Proof for gate_1__t20_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (~v1_xboole_0(k14_gate_1(X0,X1,X2,X3)) <=> (~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input t20_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (~v1_xboole_0(k14_gate_1(X0,X1,X2,X3)) <=> (~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((~(~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k14_gate_1(X0,X1,X2,X3) = k1_xboole_0) & ~(k14_gate_1(X0,X1,X2,X3) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input d14_gate_1]
14. ~! [X0,X1,X2,X3] : (~v1_xboole_0(k14_gate_1(X0,X1,X2,X3)) <=> (~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 2]
15. ! [X0,X1,X2,X3] : ((~(~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k14_gate_1(X0,X1,X2,X3) = k1_xboole_0) & ~(k14_gate_1(X0,X1,X2,X3) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 13]
16. ? [X0,X1,X2,X3] : (~v1_xboole_0(k14_gate_1(X0,X1,X2,X3)) <~> (~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2,X3] : ((k14_gate_1(X0,X1,X2,X3) = k1_xboole_0 | (~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & (k14_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 15]
20. ? [X0,X1,X2,X3] : (((v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) | v1_xboole_0(k14_gate_1(X0,X1,X2,X3))) & ((~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k14_gate_1(X0,X1,X2,X3)))) [nnf transformation 16]
21. ? [X0,X1,X2,X3] : ((v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k14_gate_1(X0,X1,X2,X3))) & ((~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k14_gate_1(X0,X1,X2,X3)))) [flattening 20]
22. ? [X0,X1,X2,X3] : ((v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k14_gate_1(X0,X1,X2,X3))) & ((~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k14_gate_1(X0,X1,X2,X3)))) => ((v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3))) & ((~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3)))) [choice axiom]
23. (v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3))) & ((~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3))) [skolemisation 21,22]
28. ~v1_xboole_0(sK0) | ~v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~v1_xboole_0(sK1) | ~v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. ~v1_xboole_0(sK2) | ~v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~v1_xboole_0(sK3) | ~v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3)) [cnf transformation 23] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
35. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
36. k14_gate_1(X0,X1,X2,X3) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:16,wCS:6400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
37. k14_gate_1(X0,X1,X2,X3) = k1_xboole_0 | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. k14_gate_1(X0,X1,X2,X3) = k1_xboole_0 | ~v1_xboole_0(X1) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. k14_gate_1(X0,X1,X2,X3) = k1_xboole_0 | ~v1_xboole_0(X2) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. k14_gate_1(X0,X1,X2,X3) = k1_xboole_0 | ~v1_xboole_0(X3) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. 1 <=> v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3)) [avatar definition]
45. ~v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3)) <- (~1) [avatar component clause 44] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
46. v1_xboole_0(k14_gate_1(sK0,sK1,sK2,sK3)) <- (1) [avatar component clause 44] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
48. 2 <=> v1_xboole_0(sK0) [avatar definition]
49. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 48] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
52. 3 <=> v1_xboole_0(sK1) [avatar definition]
53. ~v1_xboole_0(sK1) <- (~3) [avatar component clause 52] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. 4 <=> v1_xboole_0(sK2) [avatar definition]
57. ~v1_xboole_0(sK2) <- (~4) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 5 <=> v1_xboole_0(sK3) [avatar definition]
61. ~v1_xboole_0(sK3) <- (~5) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. v1_xboole_0(sK3) <- (5) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. 1 | 2 | 3 | 4 | 5 [avatar split clause 32,60,56,52,48,44]
64. ~1 | ~5 [avatar split clause 31,60,44]
65. ~1 | ~4 [avatar split clause 30,56,44]
66. ~1 | ~3 [avatar split clause 29,52,44]
67. ~1 | ~2 [avatar split clause 28,48,44]
69. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
71. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 69] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
72. 6 [avatar split clause 33,69]
84. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK0) <- (~1) [superposition 45,37] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
85. ~v1_xboole_0(sK0) <- (~1, 6) [subsumption resolution 84,71] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
88. ~2 | 1 | ~6 [avatar split clause 85,69,44,48]
89. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK1) <- (~1) [superposition 45,38] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
90. ~v1_xboole_0(sK1) <- (~1, 6) [subsumption resolution 89,71] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
93. ~3 | 1 | ~6 [avatar split clause 90,69,44,52]
94. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK2) <- (~1) [superposition 45,39] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. ~v1_xboole_0(sK2) <- (~1, 6) [subsumption resolution 94,71] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. ~4 | 1 | ~6 [avatar split clause 95,69,44,56]
99. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK3) <- (~1) [superposition 45,40] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. ~v1_xboole_0(sK3) <- (~1, 6) [subsumption resolution 99,71] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. $false <- (~1, 5, 6) [subsumption resolution 100,62] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
102. 1 | ~5 | ~6 [avatar contradiction clause 101]
115. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1) [superposition 46,36] {a:1,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~5) [subsumption resolution 115,61] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
117. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~4, ~5) [subsumption resolution 116,57] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK0) <- (1, ~3, ~4, ~5) [subsumption resolution 117,53] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (1, ~2, ~3, ~4, ~5) [subsumption resolution 118,49] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
121. 9 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
123. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (9) [avatar component clause 121] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. 9 | ~1 | 2 | 3 | 4 | 5 [avatar split clause 119,60,56,52,48,44,121]
126. ~v1_xboole_0(k1_xboole_0) <- (9) [resolution 123,35] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
128. $false <- (6, 9) [subsumption resolution 126,71] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
129. ~6 | ~9 [avatar contradiction clause 128]
130. $false [avatar sat refutation 63,64,65,66,67,72,88,93,98,102,124,129] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:79,thDist:-79}
% SZS output end Proof for gate_1__t20_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 65
% Active clauses: 21
% Passive clauses: 21
% Final active clauses: 16
% Main loop iterations started: 21
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 8
% Backward superposition: 9
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 13
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 3
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

