i: [1,64,0,5,1,1,0,0] t82_relat_1
a: [3,95,0,5,1,64]
s: 95
i: [1,65,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,100,0,2,1,65]
s: 100
i: [1,66,0,2,1,0,0,1] fc1_xboole_0
a: [3,105,0,2,1,66]
s: 105
i: [1,85,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,115,0,2,1,85]
s: 115
i: [1,86,0,2,1,0,0,2] rc2_xboole_0
a: [3,120,0,2,1,86]
s: 120
i: [1,87,0,2,1,0,0,2] rc1_xboole_0
a: [3,125,0,2,1,87]
s: 125
i: [1,89,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,130,0,2,1,89]
s: 130
i: [1,88,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,135,0,2,1,88]
s: 135
i: [1,91,0,2,1,0,0,2] rc1_relat_1
a: [3,140,0,2,1,91]
s: 140
i: [1,90,0,2,1,0,0,2] rc1_relat_1
a: [3,145,0,2,1,90]
s: 145
i: [1,67,0,3,1,0,0,255] d2_xboole_0
a: [3,110,0,3,1,67]
s: 110
r: 100
i: [1,68,0,4,2,0,0,255] cc2_zfmisc_1
s: 68
i: [1,69,0,4,2,0,0,255] cc1_zfmisc_1
s: 69
d: [2,148,1,2,1,1,40,68,125]
a: [3,153,1,2,1,148]
s: 153
i: [1,70,0,4,2,0,0,2] cc1_relat_1
s: 70
i: [1,76,0,4,1,0,0,255] existence_m1_subset_1
s: 76
d: [2,149,1,2,1,1,40,68,105]
a: [3,158,1,2,1,149]
s: 158
i: [1,71,0,5,2,0,0,1] t6_boole
s: 71
d: [2,175,1,7,2,1,44,95,71]
s: 175
i: [1,75,0,5,2,0,0,255] t7_boole
s: 75
i: [1,72,0,6,2,0,0,255] t1_subset
s: 72
d: [2,160,1,2,1,1,40,70,125]
a: [3,165,1,2,1,160]
s: 165
i: [1,77,0,6,2,0,0,1] dt_k5_relat_1
s: 77
d: [2,161,1,2,1,1,40,70,105]
a: [3,170,1,2,1,161]
s: 170
i: [1,78,0,6,2,0,0,1] dt_k5_relat_1
s: 78
i: [1,80,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 80
i: [1,74,0,7,3,0,0,1] t8_boole
s: 74
i: [1,73,0,8,3,0,0,255] t2_subset
s: 73
d: [2,177,1,6,2,0,40,73,76]
s: 177
i: [1,81,0,8,3,0,0,1] fc16_relat_1
s: 81
i: [1,82,0,8,3,0,0,1] fc16_relat_1
s: 82
d: [2,178,2,6,2,0,40,177,80]
s: 178
i: [1,83,0,8,3,0,0,1] fc17_relat_1
s: 83
d: [2,183,0,6,2,0,21,83,70]
s: 183
i: [1,84,0,8,3,0,0,1] fc17_relat_1
s: 84
i: [1,79,0,11,2,0,0,1] rd5_relat_1
s: 79
d: [2,187,2,13,3,1,44,175,79]
d: [2,191,2,6,2,1,19,187]
s: 191
d: [2,184,1,6,2,0,40,183,70]
s: 184
d: [2,192,3,4,2,1,40,191,183]
d: [2,196,3,2,1,1,21,192,70]
s: 196
r: 145
r: 135
r: 120
r: 191
r: 184
r: 183
r: 175
r: 81
r: 75
r: 74
r: 71
r: 70
r: 68
d: [2,208,4,0,0,2,40,196,105]
e: 208
d: [2,207,4,0,0,2,40,196,125]
e: 207
d: [2,206,4,0,0,2,52,125,196]
e: 206
f: 206,87,64
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t82_relat_1
% SZS output start Proof for relat_1__t82_relat_1
1. ! [X0] : k1_xboole_0 = k5_relat_1(k1_xboole_0,X0) [input t82_relat_1]
2. ~! [X0] : k1_xboole_0 = k5_relat_1(k1_xboole_0,X0) [negated conjecture 1]
12. ! [X0] : (v1_xboole_0(X0) => v1_relat_1(X0)) [input cc1_relat_1]
15. ! [X0,X1] : ((v1_relat_1(X0) & v1_xboole_0(X0)) => (v1_relat_1(k5_relat_1(X0,X1)) & v1_xboole_0(k5_relat_1(X0,X1)))) [input fc17_relat_1]
17. ? [X0] : v1_xboole_0(X0) [input rc1_xboole_0]
19. ! [X0,X1] : (v1_relat_1(X0) => k5_relat_1(X0,X1) = k5_relat_1(k5_relat_1(X0,X1),X1)) [input rd5_relat_1]
26. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
32. ? [X0] : k1_xboole_0 != k5_relat_1(k1_xboole_0,X0) [ennf transformation 2]
35. ! [X0] : (v1_relat_1(X0) | ~v1_xboole_0(X0)) [ennf transformation 12]
36. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 26]
44. ! [X0,X1] : (k5_relat_1(X0,X1) = k5_relat_1(k5_relat_1(X0,X1),X1) | ~v1_relat_1(X0)) [ennf transformation 19]
48. ! [X0,X1] : ((v1_relat_1(k5_relat_1(X0,X1)) & v1_xboole_0(k5_relat_1(X0,X1))) | (~v1_relat_1(X0) | ~v1_xboole_0(X0))) [ennf transformation 15]
49. ! [X0,X1] : ((v1_relat_1(k5_relat_1(X0,X1)) & v1_xboole_0(k5_relat_1(X0,X1))) | ~v1_relat_1(X0) | ~v1_xboole_0(X0)) [flattening 48]
50. ? [X0] : k1_xboole_0 != k5_relat_1(k1_xboole_0,X0) => k1_xboole_0 != k5_relat_1(k1_xboole_0,sK0) [choice axiom]
51. k1_xboole_0 != k5_relat_1(k1_xboole_0,sK0) [skolemisation 32,50]
58. ? [X0] : v1_xboole_0(X0) => v1_xboole_0(sK4) [choice axiom]
59. v1_xboole_0(sK4) [skolemisation 17,58]
64. k1_xboole_0 != k5_relat_1(k1_xboole_0,sK0) [cnf transformation 51] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. ~v1_xboole_0(X0) | v1_relat_1(X0) [cnf transformation 35] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
71. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 36] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. k5_relat_1(X0,X1) = k5_relat_1(k5_relat_1(X0,X1),X1) | ~v1_relat_1(X0) [cnf transformation 44] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
83. v1_xboole_0(k5_relat_1(X0,X1)) | ~v1_relat_1(X0) | ~v1_xboole_0(X0) [cnf transformation 49] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
87. v1_xboole_0(sK4) [cnf transformation 59] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
93. 1 <=> k1_xboole_0 = k5_relat_1(k1_xboole_0,sK0) [avatar definition]
95. k1_xboole_0 != k5_relat_1(k1_xboole_0,sK0) <- (~1) [avatar component clause 93] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
96. ~1 [avatar split clause 64,93]
123. 7 <=> v1_xboole_0(sK4) [avatar definition]
125. v1_xboole_0(sK4) <- (7) [avatar component clause 123] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
126. 7 [avatar split clause 87,123]
175. k5_relat_1(X0,sK0) != X0 | ~v1_xboole_0(X0) <- (~1) [superposition 95,71] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
183. v1_xboole_0(k5_relat_1(X0,X1)) | ~v1_xboole_0(X0) [subsumption resolution 83,70] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
187. k5_relat_1(X0,sK0) != k5_relat_1(X0,sK0) | ~v1_xboole_0(k5_relat_1(X0,sK0)) | ~v1_relat_1(X0) <- (~1) [superposition 175,79] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
191. ~v1_xboole_0(k5_relat_1(X0,sK0)) | ~v1_relat_1(X0) <- (~1) [trivial inequality removal 187] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
192. ~v1_relat_1(X0) | ~v1_xboole_0(X0) <- (~1) [resolution 191,183] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
196. ~v1_xboole_0(X0) <- (~1) [subsumption resolution 192,70] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
206. $false <- (~1, 7) [unit resulting resolution 125,196] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
211. 1 | ~7 [avatar contradiction clause 206]
212. $false [avatar sat refutation 96,126,211] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:76,thDist:-76}
% SZS output end Proof for relat_1__t82_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 28
% Generated clauses: 78
% Active clauses: 34
% Passive clauses: 43
% Final active clauses: 20
% Final passive clauses: 4
% Main loop iterations started: 39
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Trivial inequalities: 1
% Fw subsumption resolutions: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 13
% Backward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 16
% Unit resulting resolution: 1
% Forward superposition: 3
% Backward superposition: 7
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 11
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

