i: [1,70,0,12,4,1,0,0] t53_mcart_1
a: [3,106,0,3,1,70]
s: 106
i: [1,69,0,7,1,1,0,0] t53_mcart_1
a: [3,123,0,7,1,69]
s: 123
i: [1,71,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,133,0,2,1,71]
s: 133
i: [1,72,0,2,1,0,0,1] fc1_xboole_0
a: [3,138,0,2,1,72]
s: 138
i: [1,94,0,2,1,0,0,2] rc2_xboole_0
a: [3,148,0,2,1,94]
s: 148
i: [1,95,0,2,1,0,0,2] rc1_xboole_0
a: [3,153,0,2,1,95]
s: 153
i: [1,97,0,2,1,0,0,255] rc1_relat_1
a: [3,158,0,2,1,97]
s: 158
i: [1,96,0,2,1,0,0,255] rc1_relat_1
a: [3,163,0,2,1,96]
s: 163
i: [1,98,0,2,1,0,0,255] rc2_relat_1
a: [3,168,0,2,1,98]
s: 168
i: [1,68,0,11,1,1,0,0] t53_mcart_1
a: [3,128,0,11,1,68]
s: 128
i: [1,73,0,3,1,0,0,255] d2_xboole_0
a: [3,143,0,3,1,73]
s: 143
r: 133
i: [1,74,0,4,2,0,0,255] cc1_relat_1
s: 74
i: [1,91,0,4,1,0,0,255] existence_m1_subset_1
s: 91
d: [2,171,1,2,1,1,40,74,153]
a: [3,176,1,2,1,171]
s: 176
i: [1,75,0,5,2,0,0,1] t6_boole
s: 75
i: [1,81,0,5,2,0,0,255] t7_boole
s: 81
d: [2,172,1,2,1,1,40,74,138]
a: [3,181,1,2,1,172]
s: 181
i: [1,78,0,6,2,0,0,255] t1_subset
s: 78
i: [1,92,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 92
i: [1,80,0,7,3,0,0,1] t8_boole
s: 80
i: [1,99,0,7,1,0,0,1] t51_mcart_1
s: 99
i: [1,100,0,7,1,0,0,1] t51_mcart_1
s: 100
i: [1,101,0,7,1,0,0,1] t51_mcart_1
s: 101
i: [1,102,0,7,1,0,0,1] t51_mcart_1
s: 102
i: [1,79,0,8,3,0,0,255] t2_subset
s: 79
i: [1,93,0,14,5,0,0,1] fc12_subset_1
s: 93
d: [2,191,1,14,5,1,44,93,128]
a: [3,215,1,6,1,191]
s: 215
i: [1,82,0,19,5,0,0,1] t51_mcart_1
s: 82
d: [2,190,1,6,2,0,40,79,91]
s: 190
i: [1,87,0,26,6,0,0,1] t52_mcart_1
s: 87
i: [1,88,0,26,6,0,0,1] t52_mcart_1
s: 88
d: [2,225,1,26,6,1,44,87,128]
a: [3,240,1,3,1,225]
s: 240
r: 128
r: 106
d: [2,273,0,11,1,2,23,128,240]
d: [2,275,0,7,1,2,22,273,102]
d: [2,276,0,0,0,3,21,275,123]
e: 276
r: 240
s: 106
a: [3,239,1,3,1,225]
s: 239
s: 128
i: [1,89,0,26,6,0,0,1] t52_mcart_1
s: 89
a: [3,244,1,3,1,225]
s: 244
r: 128
d: [2,305,0,11,1,2,23,128,244]
d: [2,306,0,7,1,2,22,305,101]
d: [2,307,0,0,0,3,21,306,123]
e: 307
r: 244
s: 128
a: [3,243,1,3,1,225]
s: 243
i: [1,90,0,26,6,0,0,1] t52_mcart_1
s: 90
a: [3,248,1,3,1,225]
s: 248
r: 128
d: [2,341,0,11,1,2,23,128,248]
d: [2,342,0,7,1,2,22,341,100]
d: [2,343,0,0,0,3,21,342,123]
e: 343
r: 248
a: [3,247,1,3,1,225]
s: 247
s: 128
a: [3,252,1,3,1,225]
s: 252
r: 128
d: [2,371,0,11,1,2,23,128,252]
d: [2,372,0,7,1,2,22,371,99]
d: [2,373,0,0,0,3,21,372,123]
e: 373
r: 252
s: 128
a: [3,251,1,3,1,225]
s: 251
d: [2,401,2,0,0,6,52,106,239,243,247,128,251,87]
e: 401
r: 106
a: [3,105,0,3,1,70]
s: 105
r: 128
a: [3,110,0,3,1,70]
s: 110
d: [2,404,2,3,1,2,44,239,105]
a: [3,413,2,3,1,404]
s: 413
d: [2,403,0,11,1,2,23,128,105]
a: [3,408,0,11,1,403]
s: 408
d: [2,415,3,0,0,6,52,110,413,243,247,251,408,88]
e: 415
r: 110
a: [3,109,0,3,1,70]
s: 109
r: 408
a: [3,114,0,3,1,70]
s: 114
d: [2,446,0,11,1,2,23,408,109]
a: [3,451,0,11,1,446]
s: 451
d: [2,447,2,3,1,2,44,243,109]
a: [3,456,2,3,1,447]
s: 456
d: [2,490,3,0,0,6,52,114,413,251,247,451,456,89]
e: 490
r: 114
a: [3,113,0,3,1,70]
s: 113
r: 451
a: [3,118,0,3,1,70]
s: 118
d: [2,493,2,3,1,2,44,247,113]
a: [3,502,2,3,1,493]
s: 502
d: [2,492,0,11,1,2,23,451,113]
a: [3,497,0,11,1,492]
s: 497
d: [2,504,3,0,0,6,52,118,413,456,502,251,497,90]
e: 504
f: 504,493,492,490,447,446,415,404,403,401,373,343,307,276,68,69,70
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t53_mcart_1
% SZS output start Proof for mcart_1__t53_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7) => ((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k4_zfmisc_1(X0,X1,X2,X3) = k1_xboole_0)) [input t53_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7) => ((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k4_zfmisc_1(X0,X1,X2,X3) = k1_xboole_0)) [negated conjecture 1]
21. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((k1_xboole_0 != X3 & k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) <=> k4_zfmisc_1(X0,X1,X2,X3) != k1_xboole_0) [input t51_mcart_1]
22. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7) => ((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t52_mcart_1]
26. ~! [X0,X1,X2,X3,X4,X5,X6,X7] : (k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7) => ((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k4_zfmisc_1(X0,X1,X2,X3) = k1_xboole_0)) [flattening 2]
31. ! [X0,X1,X2,X3] : ((k1_xboole_0 != X3 & k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) <=> k4_zfmisc_1(X0,X1,X2,X3) != k1_xboole_0) [flattening 21]
32. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7) => ((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 22]
36. ? [X0,X1,X2,X3,X4,X5,X6,X7] : (((X3 != X7 | X2 != X6 | X1 != X5 | X0 != X4) & k4_zfmisc_1(X0,X1,X2,X3) != k1_xboole_0) & k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7)) [ennf transformation 26]
37. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((X3 != X7 | X2 != X6 | X1 != X5 | X0 != X4) & k4_zfmisc_1(X0,X1,X2,X3) != k1_xboole_0 & k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7)) [flattening 36]
49. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0) | k4_zfmisc_1(X0,X1,X2,X3) != k4_zfmisc_1(X4,X5,X6,X7)) [ennf transformation 32]
50. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | k4_zfmisc_1(X0,X1,X2,X3) != k4_zfmisc_1(X4,X5,X6,X7)) [flattening 49]
54. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((X3 != X7 | X2 != X6 | X1 != X5 | X0 != X4) & k4_zfmisc_1(X0,X1,X2,X3) != k1_xboole_0 & k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7)) => ((sK3 != sK7 | sK2 != sK6 | sK1 != sK5 | sK0 != sK4) & k1_xboole_0 != k4_zfmisc_1(sK0,sK1,sK2,sK3) & k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7)) [choice axiom]
55. (sK3 != sK7 | sK2 != sK6 | sK1 != sK5 | sK0 != sK4) & k1_xboole_0 != k4_zfmisc_1(sK0,sK1,sK2,sK3) & k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7) [skolemisation 37,54]
56. ! [X0,X1,X2,X3] : (((k1_xboole_0 != X3 & k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) | k4_zfmisc_1(X0,X1,X2,X3) = k1_xboole_0) & (k4_zfmisc_1(X0,X1,X2,X3) != k1_xboole_0 | (k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0))) [nnf transformation 31]
57. ! [X0,X1,X2,X3] : (((k1_xboole_0 != X3 & k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) | k4_zfmisc_1(X0,X1,X2,X3) = k1_xboole_0) & (k4_zfmisc_1(X0,X1,X2,X3) != k1_xboole_0 | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 56]
68. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7) [cnf transformation 55] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. k1_xboole_0 != k4_zfmisc_1(sK0,sK1,sK2,sK3) [cnf transformation 55] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. sK3 != sK7 | sK2 != sK6 | sK1 != sK5 | sK0 != sK4 [cnf transformation 55] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
83. k1_xboole_0 != X0 | k4_zfmisc_1(X0,X1,X2,X3) = k1_xboole_0 [cnf transformation 57] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:1,thDist:-1}
84. k1_xboole_0 != X1 | k4_zfmisc_1(X0,X1,X2,X3) = k1_xboole_0 [cnf transformation 57] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:1,thDist:-1}
85. k1_xboole_0 != X2 | k4_zfmisc_1(X0,X1,X2,X3) = k1_xboole_0 [cnf transformation 57] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:1,thDist:-1}
86. k1_xboole_0 != X3 | k4_zfmisc_1(X0,X1,X2,X3) = k1_xboole_0 [cnf transformation 57] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:1,thDist:-1}
87. k4_zfmisc_1(X0,X1,X2,X3) != k4_zfmisc_1(X4,X5,X6,X7) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X0 = X4 [cnf transformation 50] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
88. k4_zfmisc_1(X0,X1,X2,X3) != k4_zfmisc_1(X4,X5,X6,X7) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X1 = X5 [cnf transformation 50] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
89. k4_zfmisc_1(X0,X1,X2,X3) != k4_zfmisc_1(X4,X5,X6,X7) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X2 = X6 [cnf transformation 50] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
90. k4_zfmisc_1(X0,X1,X2,X3) != k4_zfmisc_1(X4,X5,X6,X7) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X3 = X7 [cnf transformation 50] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
99. k1_xboole_0 = k4_zfmisc_1(X0,X1,X2,k1_xboole_0) [equality resolution 86] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
100. k1_xboole_0 = k4_zfmisc_1(X0,X1,k1_xboole_0,X3) [equality resolution 85] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. k1_xboole_0 = k4_zfmisc_1(X0,k1_xboole_0,X2,X3) [equality resolution 84] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
102. k1_xboole_0 = k4_zfmisc_1(k1_xboole_0,X1,X2,X3) [equality resolution 83] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. 1 <=> sK0 = sK4 [avatar definition]
105. sK0 = sK4 <- (1) [avatar component clause 104] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
106. sK0 != sK4 <- (~1) [avatar component clause 104] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
108. 2 <=> sK1 = sK5 [avatar definition]
109. sK1 = sK5 <- (2) [avatar component clause 108] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
110. sK1 != sK5 <- (~2) [avatar component clause 108] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
112. 3 <=> sK2 = sK6 [avatar definition]
113. sK2 = sK6 <- (3) [avatar component clause 112] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
114. sK2 != sK6 <- (~3) [avatar component clause 112] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. 4 <=> sK3 = sK7 [avatar definition]
118. sK3 != sK7 <- (~4) [avatar component clause 116] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
119. ~1 | ~2 | ~3 | ~4 [avatar split clause 70,116,112,108,104]
121. 5 <=> k1_xboole_0 = k4_zfmisc_1(sK0,sK1,sK2,sK3) [avatar definition]
123. k1_xboole_0 != k4_zfmisc_1(sK0,sK1,sK2,sK3) <- (~5) [avatar component clause 121] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
124. ~5 [avatar split clause 69,121]
126. 6 <=> k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7) [avatar definition]
128. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7) <- (6) [avatar component clause 126] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
129. 6 [avatar split clause 68,126]
238. 22 <=> k1_xboole_0 = sK4 [avatar definition]
239. k1_xboole_0 != sK4 <- (~22) [avatar component clause 238] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
240. k1_xboole_0 = sK4 <- (22) [avatar component clause 238] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
242. 23 <=> k1_xboole_0 = sK5 [avatar definition]
243. k1_xboole_0 != sK5 <- (~23) [avatar component clause 242] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
244. k1_xboole_0 = sK5 <- (23) [avatar component clause 242] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
246. 24 <=> k1_xboole_0 = sK6 [avatar definition]
247. k1_xboole_0 != sK6 <- (~24) [avatar component clause 246] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
248. k1_xboole_0 = sK6 <- (24) [avatar component clause 246] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
250. 25 <=> k1_xboole_0 = sK7 [avatar definition]
251. k1_xboole_0 != sK7 <- (~25) [avatar component clause 250] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
252. k1_xboole_0 = sK7 <- (25) [avatar component clause 250] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
273. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(k1_xboole_0,sK5,sK6,sK7) <- (6, 22) [backward demodulation 128,240] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
275. k1_xboole_0 = k4_zfmisc_1(sK0,sK1,sK2,sK3) <- (6, 22) [forward demodulation 273,102] {a:0,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
276. $false <- (~5, 6, 22) [subsumption resolution 275,123] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
277. 5 | ~6 | ~22 [avatar contradiction clause 276]
305. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,k1_xboole_0,sK6,sK7) <- (6, 23) [backward demodulation 128,244] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
306. k1_xboole_0 = k4_zfmisc_1(sK0,sK1,sK2,sK3) <- (6, 23) [forward demodulation 305,101] {a:0,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
307. $false <- (~5, 6, 23) [subsumption resolution 306,123] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
308. 5 | ~6 | ~23 [avatar contradiction clause 307]
341. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,k1_xboole_0,sK7) <- (6, 24) [backward demodulation 128,248] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
342. k1_xboole_0 = k4_zfmisc_1(sK0,sK1,sK2,sK3) <- (6, 24) [forward demodulation 341,100] {a:0,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
343. $false <- (~5, 6, 24) [subsumption resolution 342,123] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
344. 5 | ~6 | ~24 [avatar contradiction clause 343]
371. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,k1_xboole_0) <- (6, 25) [backward demodulation 128,252] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
372. k1_xboole_0 = k4_zfmisc_1(sK0,sK1,sK2,sK3) <- (6, 25) [forward demodulation 371,99] {a:0,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
373. $false <- (~5, 6, 25) [subsumption resolution 372,123] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
374. 5 | ~6 | ~25 [avatar contradiction clause 373]
401. $false <- (~1, 6, ~22, ~23, ~24, ~25) [unit resulting resolution 106,239,243,247,128,251,87] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:17,thDist:-17}
402. 1 | ~6 | 22 | 23 | 24 | 25 [avatar contradiction clause 401]
403. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK0,sK5,sK6,sK7) <- (1, 6) [backward demodulation 128,105] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
404. k1_xboole_0 != sK0 <- (1, ~22) [superposition 239,105] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
406. 30 <=> k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK0,sK5,sK6,sK7) [avatar definition]
408. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK0,sK5,sK6,sK7) <- (30) [avatar component clause 406] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
409. 30 | ~1 | ~6 [avatar split clause 403,126,104,406]
411. 31 <=> k1_xboole_0 = sK0 [avatar definition]
413. k1_xboole_0 != sK0 <- (~31) [avatar component clause 411] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
414. ~31 | ~1 | 22 [avatar split clause 404,238,104,411]
415. $false <- (~2, ~23, ~24, ~25, 30, ~31) [unit resulting resolution 110,413,243,247,251,408,88] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
426. 2 | 23 | 24 | 25 | ~30 | 31 [avatar contradiction clause 415]
446. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK0,sK1,sK6,sK7) <- (2, 30) [backward demodulation 408,109] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
447. k1_xboole_0 != sK1 <- (2, ~23) [superposition 243,109] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
449. 32 <=> k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK0,sK1,sK6,sK7) [avatar definition]
451. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK0,sK1,sK6,sK7) <- (32) [avatar component clause 449] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
452. 32 | ~2 | ~30 [avatar split clause 446,406,108,449]
454. 33 <=> k1_xboole_0 = sK1 [avatar definition]
456. k1_xboole_0 != sK1 <- (~33) [avatar component clause 454] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
457. ~33 | ~2 | 23 [avatar split clause 447,242,108,454]
490. $false <- (~3, ~24, ~25, ~31, 32, ~33) [unit resulting resolution 114,413,251,247,451,456,89] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:21,thDist:-21}
491. 3 | 24 | 25 | 31 | ~32 | 33 [avatar contradiction clause 490]
492. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK0,sK1,sK2,sK7) <- (3, 32) [backward demodulation 451,113] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
493. k1_xboole_0 != sK2 <- (3, ~24) [superposition 247,113] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
495. 36 <=> k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK0,sK1,sK2,sK7) [avatar definition]
497. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK0,sK1,sK2,sK7) <- (36) [avatar component clause 495] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
498. 36 | ~3 | ~32 [avatar split clause 492,449,112,495]
500. 37 <=> k1_xboole_0 = sK2 [avatar definition]
502. k1_xboole_0 != sK2 <- (~37) [avatar component clause 500] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
503. ~37 | ~3 | 24 [avatar split clause 493,246,112,500]
504. $false <- (~4, ~25, ~31, ~33, 36, ~37) [unit resulting resolution 118,413,456,502,251,497,90] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
515. 4 | 25 | 31 | 33 | ~36 | 37 [avatar contradiction clause 504]
532. $false [avatar sat refutation 119,124,129,277,308,344,374,402,409,414,426,452,457,491,498,503,515] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:394,thDist:-394}
% SZS output end Proof for mcart_1__t53_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 31
% Generated clauses: 298
% Active clauses: 58
% Passive clauses: 93
% Final active clauses: 42
% Final passive clauses: 35
% Main loop iterations started: 58
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 80
% Fw demodulations: 4
% Bw demodulations: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Equational tautologies: 4
% Forward subsumptions: 40
% 
% >>> Generating Inferences
% Binary resolution: 6
% Unit resulting resolution: 4
% Forward superposition: 35
% Backward superposition: 68
% Self superposition: 1
% Equality resolution: 4
% 
% >>> AVATAR
% Split clauses: 23
% Split components: 70
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 54
% SAT solver unit clauses: 10
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 5245
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.002 s
% forward demodulation: 0.001 s
% unit resulting resolution: 0.002 s
% other: 0.007 s ( own 0.002 s ) 

