i: [1,56,0,2,1,1,0,0] t59_xxreal_2
a: [3,95,0,2,1,56]
s: 95
i: [1,55,0,2,1,1,0,0] t59_xxreal_2
a: [3,100,0,2,1,55]
s: 100
i: [1,57,0,3,1,1,0,0] t59_xxreal_2
a: [3,90,0,3,1,57]
s: 90
i: [1,58,0,5,1,1,0,0] t59_xxreal_2
a: [3,85,0,5,1,58]
s: 85
i: [1,79,0,2,1,0,0,2] rc1_xxreal_0
a: [3,112,0,2,1,79]
s: 112
i: [1,59,0,3,1,0,0,255] fc1_finsub_1
s: 59
i: [1,76,0,3,1,0,0,1] reflexivity_r1_tarski
s: 76
i: [1,60,0,4,2,0,0,255] cc1_finsub_1
s: 60
d: [2,114,1,3,1,0,40,60,59]
s: 114
i: [1,61,0,4,2,0,0,255] cc1_finsub_1
s: 61
i: [1,75,0,4,1,0,0,255] existence_m1_subset_1
s: 75
d: [2,115,1,3,1,0,40,61,59]
s: 115
i: [1,62,0,5,2,0,0,1] dt_k1_xxreal_2
s: 62
i: [1,77,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,107,0,5,2,77]
s: 107
i: [1,71,0,6,2,0,0,2] existence_m1_xxreal_2
s: 71
i: [1,72,0,6,3,0,0,255] cc2_finsub_1
s: 72
i: [1,68,0,7,3,0,0,255] cc13_membered
s: 68
i: [1,69,0,7,3,0,0,2] dt_m1_xxreal_2
s: 69
d: [2,117,1,5,2,0,40,68,75]
s: 117
i: [1,73,0,7,2,0,0,255] t3_subset
s: 73
i: [1,74,0,7,2,0,0,255] t3_subset
s: 74
d: [2,118,1,7,3,0,40,69,71]
d: [2,119,1,5,2,0,18,118]
s: 119
i: [1,70,0,8,3,0,0,255] cc20_membered
s: 70
d: [2,120,1,5,1,0,40,73,75]
s: 120
i: [1,81,0,9,3,0,0,255] d3_xxreal_2
s: 81
d: [2,125,0,6,2,0,21,81,62]
s: 125
d: [2,123,1,6,2,0,40,70,75]
s: 123
i: [1,78,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 78
d: [2,128,1,11,3,1,40,78,85]
a: [3,134,1,3,1,128]
s: 134
d: [2,144,2,0,0,2,52,95,134,62]
e: 144
d: [2,145,2,2,1,1,40,134,62]
d: [2,147,2,0,0,2,21,145,95]
e: 147
r: 134
i: [1,80,0,12,4,0,0,255] d3_xxreal_2
s: 80
d: [2,149,0,9,3,0,21,80,62]
s: 149
a: [3,133,1,3,1,128]
s: 133
i: [1,63,0,13,5,0,0,1] t6_xxreal_2
s: 63
a: [3,138,1,3,1,128]
s: 138
d: [2,161,2,0,0,2,52,100,138,62]
e: 161
d: [2,162,2,2,1,1,40,138,62]
d: [2,164,2,0,0,2,21,162,100]
e: 164
r: 138
i: [1,66,0,16,5,0,0,255] d3_xxreal_2
s: 66
d: [2,166,0,14,4,0,21,66,69]
s: 166
a: [3,137,1,3,1,128]
s: 137
i: [1,67,0,16,5,0,0,255] d3_xxreal_2
s: 67
d: [2,150,1,6,2,1,40,149,85]
d: [2,151,1,4,1,2,21,150,100]
a: [3,155,1,4,1,151]
s: 155
d: [2,171,0,14,4,0,21,67,69]
s: 171
a: [3,142,1,5,1,128]
s: 142
d: [2,124,1,7,3,0,40,70,74]
s: 124
d: [2,122,1,8,3,0,40,74,68]
s: 122
d: [2,178,2,5,2,1,40,122,90]
a: [3,184,2,3,1,178]
s: 184
d: [2,158,1,13,5,0,40,63,125]
d: [2,159,1,11,4,0,18,158]
s: 159
d: [2,190,1,9,3,0,21,159,124]
s: 190
d: [2,191,2,0,0,3,52,95,90,155,190]
e: 191
d: [2,192,2,5,2,1,40,190,155]
d: [2,196,2,2,1,2,21,192,90]
d: [2,197,2,0,0,3,21,196,95]
e: 197
f: 197,151,55,56,57,58
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_2__t59_xxreal_2
% SZS output start Proof for xxreal_2__t59_xxreal_2
1. ! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => (r1_tarski(X0,X1) => r1_xxreal_0(k1_xxreal_2(X0),k1_xxreal_2(X1))))) [input t59_xxreal_2]
2. ~! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => (r1_tarski(X0,X1) => r1_xxreal_0(k1_xxreal_2(X0),k1_xxreal_2(X1))))) [negated conjecture 1]
5. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v2_membered(X1))) [input cc20_membered]
8. ! [X0] : (v2_membered(X0) => ! [X1] : (v1_xxreal_0(X1) => (k1_xxreal_2(X0) = X1 <=> (! [X2] : (m1_xxreal_2(X2,X0) => r1_xxreal_0(X1,X2)) & m1_xxreal_2(X1,X0))))) [input d3_xxreal_2]
9. ! [X0] : (v2_membered(X0) => v1_xxreal_0(k1_xxreal_2(X0))) [input dt_k1_xxreal_2]
19. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
20. ! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => (r1_tarski(X0,X1) => ! [X2] : (m1_xxreal_2(X2,X1) => m1_xxreal_2(X2,X0))))) [input t6_xxreal_2]
21. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 19]
23. ? [X0] : (? [X1] : ((~r1_xxreal_0(k1_xxreal_2(X0),k1_xxreal_2(X1)) & r1_tarski(X0,X1)) & v2_membered(X1)) & v2_membered(X0)) [ennf transformation 2]
24. ? [X0] : (? [X1] : (~r1_xxreal_0(k1_xxreal_2(X0),k1_xxreal_2(X1)) & r1_tarski(X0,X1) & v2_membered(X1)) & v2_membered(X0)) [flattening 23]
26. ! [X0] : (v1_xxreal_0(k1_xxreal_2(X0)) | ~v2_membered(X0)) [ennf transformation 9]
27. ! [X0] : (! [X1] : ((! [X2] : (m1_xxreal_2(X2,X0) | ~m1_xxreal_2(X2,X1)) | ~r1_tarski(X0,X1)) | ~v2_membered(X1)) | ~v2_membered(X0)) [ennf transformation 20]
28. ! [X0] : (! [X1] : (! [X2] : (m1_xxreal_2(X2,X0) | ~m1_xxreal_2(X2,X1)) | ~r1_tarski(X0,X1) | ~v2_membered(X1)) | ~v2_membered(X0)) [flattening 27]
29. ! [X0] : (! [X1] : ((k1_xxreal_2(X0) = X1 <=> (! [X2] : (r1_xxreal_0(X1,X2) | ~m1_xxreal_2(X2,X0)) & m1_xxreal_2(X1,X0))) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [ennf transformation 8]
32. ! [X0] : (! [X1] : (v2_membered(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v2_membered(X0)) [ennf transformation 5]
40. ? [X0] : (? [X1] : (~r1_xxreal_0(k1_xxreal_2(X0),k1_xxreal_2(X1)) & r1_tarski(X0,X1) & v2_membered(X1)) & v2_membered(X0)) => (? [X1] : (~r1_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(X1)) & r1_tarski(sK0,X1) & v2_membered(X1)) & v2_membered(sK0)) [choice axiom]
41. ? [X1] : (~r1_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(X1)) & r1_tarski(sK0,X1) & v2_membered(X1)) => (~r1_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1)) & r1_tarski(sK0,sK1) & v2_membered(sK1)) [choice axiom]
42. (~r1_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1)) & r1_tarski(sK0,sK1) & v2_membered(sK1)) & v2_membered(sK0) [skolemisation 24,41,40]
43. ! [X0] : (! [X1] : (((k1_xxreal_2(X0) = X1 | (? [X2] : (~r1_xxreal_0(X1,X2) & m1_xxreal_2(X2,X0)) | ~m1_xxreal_2(X1,X0))) & ((! [X2] : (r1_xxreal_0(X1,X2) | ~m1_xxreal_2(X2,X0)) & m1_xxreal_2(X1,X0)) | k1_xxreal_2(X0) != X1)) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [nnf transformation 29]
44. ! [X0] : (! [X1] : (((k1_xxreal_2(X0) = X1 | ? [X2] : (~r1_xxreal_0(X1,X2) & m1_xxreal_2(X2,X0)) | ~m1_xxreal_2(X1,X0)) & ((! [X2] : (r1_xxreal_0(X1,X2) | ~m1_xxreal_2(X2,X0)) & m1_xxreal_2(X1,X0)) | k1_xxreal_2(X0) != X1)) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [flattening 43]
45. ! [X0] : (! [X1] : (((k1_xxreal_2(X0) = X1 | ? [X2] : (~r1_xxreal_0(X1,X2) & m1_xxreal_2(X2,X0)) | ~m1_xxreal_2(X1,X0)) & ((! [X3] : (r1_xxreal_0(X1,X3) | ~m1_xxreal_2(X3,X0)) & m1_xxreal_2(X1,X0)) | k1_xxreal_2(X0) != X1)) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [rectify 44]
46. ! [X1,X0] : (? [X2] : (~r1_xxreal_0(X1,X2) & m1_xxreal_2(X2,X0)) => (~r1_xxreal_0(X1,sK2(X0,X1)) & m1_xxreal_2(sK2(X0,X1),X0))) [choice axiom]
47. ! [X0] : (! [X1] : (((k1_xxreal_2(X0) = X1 | (~r1_xxreal_0(X1,sK2(X0,X1)) & m1_xxreal_2(sK2(X0,X1),X0)) | ~m1_xxreal_2(X1,X0)) & ((! [X3] : (r1_xxreal_0(X1,X3) | ~m1_xxreal_2(X3,X0)) & m1_xxreal_2(X1,X0)) | k1_xxreal_2(X0) != X1)) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [skolemisation 45,46]
50. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 21]
55. v2_membered(sK0) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. v2_membered(sK1) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
57. r1_tarski(sK0,sK1) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. ~r1_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1)) [cnf transformation 42] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
62. v1_xxreal_0(k1_xxreal_2(X0)) | ~v2_membered(X0) [cnf transformation 26] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. ~m1_xxreal_2(X2,X1) | m1_xxreal_2(X2,X0) | ~r1_tarski(X0,X1) | ~v2_membered(X1) | ~v2_membered(X0) [cnf transformation 28] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. m1_xxreal_2(X1,X0) | k1_xxreal_2(X0) != X1 | ~v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 47] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
65. r1_xxreal_0(X1,X3) | ~m1_xxreal_2(X3,X0) | k1_xxreal_2(X0) != X1 | ~v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 47] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:2,thDist:-2}
70. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v2_membered(X1) | ~v2_membered(X0) [cnf transformation 32] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
74. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 50] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
80. r1_xxreal_0(k1_xxreal_2(X0),X3) | ~m1_xxreal_2(X3,X0) | ~v1_xxreal_0(k1_xxreal_2(X0)) | ~v2_membered(X0) [equality resolution 65] {a:0,w:12,wCS:4800,sine:255,thAx:0,allAx:2,thDist:-2}
81. m1_xxreal_2(k1_xxreal_2(X0),X0) | ~v1_xxreal_0(k1_xxreal_2(X0)) | ~v2_membered(X0) [equality resolution 64] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:2,thDist:-2}
83. 1 <=> r1_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1)) [avatar definition]
85. ~r1_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1)) <- (~1) [avatar component clause 83] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. ~1 [avatar split clause 58,83]
88. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
90. r1_tarski(sK0,sK1) <- (2) [avatar component clause 88] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. 2 [avatar split clause 57,88]
93. 3 <=> v2_membered(sK1) [avatar definition]
95. v2_membered(sK1) <- (3) [avatar component clause 93] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. 3 [avatar split clause 56,93]
98. 4 <=> v2_membered(sK0) [avatar definition]
100. v2_membered(sK0) <- (4) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. 4 [avatar split clause 55,98]
124. ~r1_tarski(X1,X2) | ~v2_membered(X2) | v2_membered(X1) [resolution 70,74] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
125. m1_xxreal_2(k1_xxreal_2(X0),X0) | ~v2_membered(X0) [subsumption resolution 81,62] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
149. r1_xxreal_0(k1_xxreal_2(X0),X3) | ~m1_xxreal_2(X3,X0) | ~v2_membered(X0) [subsumption resolution 80,62] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
150. ~m1_xxreal_2(k1_xxreal_2(sK1),sK0) | ~v2_membered(sK0) <- (~1) [resolution 149,85] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
151. ~m1_xxreal_2(k1_xxreal_2(sK1),sK0) <- (~1, 4) [subsumption resolution 150,100] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
153. 11 <=> m1_xxreal_2(k1_xxreal_2(sK1),sK0) [avatar definition]
155. ~m1_xxreal_2(k1_xxreal_2(sK1),sK0) <- (~11) [avatar component clause 153] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
156. ~11 | 1 | ~4 [avatar split clause 151,98,83,153]
158. m1_xxreal_2(k1_xxreal_2(X2),X3) | ~r1_tarski(X3,X2) | ~v2_membered(X2) | ~v2_membered(X3) | ~v2_membered(X2) [resolution 63,125] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:3,thDist:-3}
159. m1_xxreal_2(k1_xxreal_2(X2),X3) | ~r1_tarski(X3,X2) | ~v2_membered(X2) | ~v2_membered(X3) [duplicate literal removal 158] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:3,thDist:-3}
190. m1_xxreal_2(k1_xxreal_2(X2),X3) | ~r1_tarski(X3,X2) | ~v2_membered(X2) [subsumption resolution 159,124] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
192. ~r1_tarski(sK0,sK1) | ~v2_membered(sK1) <- (~11) [resolution 190,155] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
196. ~v2_membered(sK1) <- (2, ~11) [subsumption resolution 192,90] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
197. $false <- (2, 3, ~11) [subsumption resolution 196,95] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
198. ~2 | ~3 | 11 [avatar contradiction clause 197]
200. $false [avatar sat refutation 86,91,96,101,156,198] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:94,thDist:-94}
% SZS output end Proof for xxreal_2__t59_xxreal_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 19
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 25
% Generated clauses: 91
% Active clauses: 41
% Passive clauses: 52
% Final active clauses: 39
% Final passive clauses: 6
% Main loop iterations started: 46
% 
% >>> Simplifying Inferences
% Duplicate literals: 7
% Fw subsumption resolutions: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 30
% Unit resulting resolution: 3
% Factoring: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 5
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% SAT solver time: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

