i: [1,90,0,2,1,1,0,0] t79_semi_af1
a: [3,150,0,2,1,90]
s: 150
i: [1,89,0,2,1,1,0,0] t79_semi_af1
a: [3,155,0,2,1,89]
s: 155
i: [1,88,0,2,1,1,0,0] t79_semi_af1
a: [3,160,0,2,1,88]
s: 160
i: [1,94,0,3,1,1,0,0] t79_semi_af1
a: [3,130,0,3,1,94]
s: 130
i: [1,92,0,4,1,1,0,0] t79_semi_af1
a: [3,140,0,4,1,92]
s: 140
i: [1,91,0,4,1,1,0,0] t79_semi_af1
a: [3,145,0,4,1,91]
s: 145
i: [1,93,0,7,1,1,0,0] t79_semi_af1
a: [3,135,0,7,1,93]
s: 135
i: [1,121,0,2,1,0,0,2] existence_l1_struct_0
a: [3,165,0,2,1,121]
s: 165
i: [1,122,0,2,1,0,0,1] existence_l1_analoaf
a: [3,170,0,2,1,122]
s: 170
i: [1,125,0,2,1,0,0,1] rc1_semi_af1
a: [3,175,0,2,1,125]
s: 175
i: [1,124,0,2,1,0,0,1] rc1_semi_af1
a: [3,180,0,2,1,124]
s: 180
i: [1,123,0,2,1,0,0,1] rc1_semi_af1
a: [3,185,0,2,1,123]
s: 185
i: [1,109,0,4,2,0,0,1] dt_l1_analoaf
s: 109
d: [2,187,1,2,1,1,40,109,150]
a: [3,193,1,2,1,187]
s: 193
i: [1,119,0,4,1,0,0,1] existence_m1_subset_1
s: 119
i: [1,95,0,6,3,0,0,255] cc7_struct_0
s: 95
d: [2,188,1,2,1,1,40,109,170]
a: [3,198,1,2,1,188]
s: 198
i: [1,96,0,6,3,0,0,255] cc5_struct_0
s: 96
d: [2,189,1,2,1,1,40,109,185]
a: [3,203,1,2,1,189]
s: 203
i: [1,98,0,6,3,0,0,255] cc17_struct_0
s: 98
i: [1,99,0,6,3,0,0,255] cc2_struct_0
s: 99
d: [2,205,1,4,2,1,40,96,165]
a: [3,211,1,2,1,205]
s: 211
i: [1,100,0,6,3,0,0,255] cc6_struct_0
s: 100
d: [2,207,2,4,2,1,40,96,198]
a: [3,220,2,2,1,207]
s: 220
i: [1,101,0,6,3,0,0,255] cc1_struct_0
s: 101
i: [1,104,0,6,3,0,0,255] cc4_struct_0
s: 104
i: [1,102,0,7,3,0,0,255] cc8_struct_0
s: 102
i: [1,105,0,7,3,0,0,255] cc9_struct_0
s: 105
i: [1,106,0,7,3,0,0,255] cc11_struct_0
s: 106
i: [1,107,0,7,3,0,0,255] cc11_struct_0
s: 107
i: [1,110,0,7,3,0,0,255] fc9_struct_0
s: 110
i: [1,111,0,7,3,0,0,255] fc6_struct_0
s: 111
i: [1,115,0,7,3,0,0,255] fc2_struct_0
s: 115
i: [1,116,0,7,3,0,0,255] fc8_struct_0
s: 116
i: [1,117,0,7,3,0,0,255] fc7_struct_0
s: 117
i: [1,118,0,7,3,0,0,255] fc1_struct_0
s: 118
i: [1,108,0,9,4,0,0,255] cc10_struct_0
s: 108
i: [1,112,0,23,7,0,0,255] t56_semi_af1
s: 112
i: [1,120,0,26,7,0,0,1] dt_k1_semi_af1
s: 120
i: [1,114,0,35,9,0,0,255] d5_semi_af1
s: 114
i: [1,126,0,36,8,0,0,255] d5_semi_af1
s: 126
d: [2,255,1,31,9,1,44,114,135]
d: [2,262,1,27,8,1,18,255]
d: [2,263,1,23,7,2,21,262,140]
d: [2,264,1,19,6,3,21,263,145]
d: [2,265,1,17,5,4,21,264,150]
d: [2,266,1,15,4,5,21,265,155]
d: [2,267,1,13,3,6,21,266,160]
s: 267
d: [2,273,0,28,7,0,21,126,120]
s: 273
d: [2,275,1,24,7,1,44,273,135]
d: [2,278,1,20,6,1,18,275]
d: [2,280,1,16,5,2,21,278,140]
d: [2,281,1,12,4,3,21,280,145]
d: [2,282,1,10,3,4,21,281,150]
d: [2,283,1,8,2,5,21,282,155]
d: [2,284,1,6,1,6,21,283,160]
a: [3,288,1,6,1,284]
s: 288
d: [2,290,2,0,0,7,52,160,155,150,130,140,145,288,112]
e: 290
d: [2,292,2,17,6,1,40,288,112]
d: [2,294,2,14,5,2,21,292,130]
d: [2,295,2,10,4,3,21,294,145]
d: [2,296,2,6,3,4,21,295,140]
d: [2,297,2,4,2,5,21,296,150]
d: [2,298,2,2,1,6,21,297,155]
d: [2,299,2,0,0,7,21,298,160]
e: 299
f: 299,284,88,89,90,91,92,93,94
% Refutation found. Thanks to Tanya!
% SZS status Theorem for semi_af1__t79_semi_af1
% SZS output start Proof for semi_af1__t79_semi_af1
1. ! [X0] : ((l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_semi_af1(X0,X1,X1,X2) = X2 => X1 = X2)))) [input t79_semi_af1]
2. ~! [X0] : ((l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_semi_af1(X0,X1,X1,X2) = X2 => X1 = X2)))) [negated conjecture 1]
14. ! [X0] : ((l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (k1_semi_af1(X0,X1,X2,X3) = X4 <=> r3_semi_af1(X0,X3,X1,X2,X4))))))) [input d5_semi_af1]
15. ! [X0,X1,X2,X3] : ((m1_subset_1(X3,u1_struct_0(X0)) & m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => m1_subset_1(k1_semi_af1(X0,X1,X2,X3),u1_struct_0(X0))) [input dt_k1_semi_af1]
31. ! [X0] : ((l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r3_semi_af1(X0,X1,X2,X2,X1) => X1 = X2)))) [input t56_semi_af1]
32. ? [X0] : (? [X1] : (? [X2] : ((X1 != X2 & k1_semi_af1(X0,X1,X1,X2) = X2) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
33. ? [X0] : (? [X1] : (? [X2] : (X1 != X2 & k1_semi_af1(X0,X1,X1,X2) = X2 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) [flattening 32]
61. ! [X0] : (! [X1] : (! [X2] : ((X1 = X2 | ~r3_semi_af1(X0,X1,X2,X2,X1)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0))) [ennf transformation 31]
62. ! [X0] : (! [X1] : (! [X2] : (X1 = X2 | ~r3_semi_af1(X0,X1,X2,X2,X1) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0)) [flattening 61]
63. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((k1_semi_af1(X0,X1,X2,X3) = X4 <=> r3_semi_af1(X0,X3,X1,X2,X4)) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0))) [ennf transformation 14]
64. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((k1_semi_af1(X0,X1,X2,X3) = X4 <=> r3_semi_af1(X0,X3,X1,X2,X4)) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0)) [flattening 63]
73. ! [X0,X1,X2,X3] : (m1_subset_1(k1_semi_af1(X0,X1,X2,X3),u1_struct_0(X0)) | (~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0))) [ennf transformation 15]
74. ! [X0,X1,X2,X3] : (m1_subset_1(k1_semi_af1(X0,X1,X2,X3),u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0)) [flattening 73]
75. ? [X0] : (? [X1] : (? [X2] : (X1 != X2 & k1_semi_af1(X0,X1,X1,X2) = X2 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (X1 != X2 & k1_semi_af1(sK0,X1,X1,X2) = X2 & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_semi_af1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
76. ? [X1] : (? [X2] : (X1 != X2 & k1_semi_af1(sK0,X1,X1,X2) = X2 & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (sK1 != X2 & k1_semi_af1(sK0,sK1,sK1,X2) = X2 & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
77. ? [X2] : (sK1 != X2 & k1_semi_af1(sK0,sK1,sK1,X2) = X2 & m1_subset_1(X2,u1_struct_0(sK0))) => (sK1 != sK2 & sK2 = k1_semi_af1(sK0,sK1,sK1,sK2) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
78. ((sK1 != sK2 & sK2 = k1_semi_af1(sK0,sK1,sK1,sK2) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_semi_af1(sK0) & ~v2_struct_0(sK0) [skolemisation 33,77,76,75]
79. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((k1_semi_af1(X0,X1,X2,X3) = X4 | ~r3_semi_af1(X0,X3,X1,X2,X4)) & (r3_semi_af1(X0,X3,X1,X2,X4) | k1_semi_af1(X0,X1,X2,X3) != X4)) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0)) [nnf transformation 64]
88. ~v2_struct_0(sK0) [cnf transformation 78] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
89. v1_semi_af1(sK0) [cnf transformation 78] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
90. l1_analoaf(sK0) [cnf transformation 78] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
91. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 78] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
92. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 78] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
93. sK2 = k1_semi_af1(sK0,sK1,sK1,sK2) [cnf transformation 78] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
94. sK1 != sK2 [cnf transformation 78] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
112. ~r3_semi_af1(X0,X1,X2,X2,X1) | X1 = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0) [cnf transformation 62] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
113. r3_semi_af1(X0,X3,X1,X2,X4) | k1_semi_af1(X0,X1,X2,X3) != X4 | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0) [cnf transformation 79] {a:0,w:35,wCS:14000,sine:255,thAx:0,allAx:1,thDist:-1}
120. m1_subset_1(k1_semi_af1(X0,X1,X2,X3),u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0) [cnf transformation 74] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
126. r3_semi_af1(X0,X3,X1,X2,k1_semi_af1(X0,X1,X2,X3)) | ~m1_subset_1(k1_semi_af1(X0,X1,X2,X3),u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0) [equality resolution 113] {a:0,w:36,wCS:14400,sine:255,thAx:0,allAx:1,thDist:-1}
128. 1 <=> sK1 = sK2 [avatar definition]
130. sK1 != sK2 <- (~1) [avatar component clause 128] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
131. ~1 [avatar split clause 94,128]
133. 2 <=> sK2 = k1_semi_af1(sK0,sK1,sK1,sK2) [avatar definition]
135. sK2 = k1_semi_af1(sK0,sK1,sK1,sK2) <- (2) [avatar component clause 133] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
136. 2 [avatar split clause 93,133]
138. 3 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
140. m1_subset_1(sK2,u1_struct_0(sK0)) <- (3) [avatar component clause 138] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
141. 3 [avatar split clause 92,138]
143. 4 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
145. m1_subset_1(sK1,u1_struct_0(sK0)) <- (4) [avatar component clause 143] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
146. 4 [avatar split clause 91,143]
148. 5 <=> l1_analoaf(sK0) [avatar definition]
150. l1_analoaf(sK0) <- (5) [avatar component clause 148] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
151. 5 [avatar split clause 90,148]
153. 6 <=> v1_semi_af1(sK0) [avatar definition]
155. v1_semi_af1(sK0) <- (6) [avatar component clause 153] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
156. 6 [avatar split clause 89,153]
158. 7 <=> v2_struct_0(sK0) [avatar definition]
160. ~v2_struct_0(sK0) <- (~7) [avatar component clause 158] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
161. ~7 [avatar split clause 88,158]
273. r3_semi_af1(X0,X3,X1,X2,k1_semi_af1(X0,X1,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0) [subsumption resolution 126,120] {a:0,w:28,wCS:11200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
275. r3_semi_af1(sK0,sK2,sK1,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (2) [superposition 273,135] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
278. r3_semi_af1(sK0,sK2,sK1,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (2) [duplicate literal removal 275] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
280. r3_semi_af1(sK0,sK2,sK1,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (2, 3) [subsumption resolution 278,140] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
281. r3_semi_af1(sK0,sK2,sK1,sK1,sK2) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (2, 3, 4) [subsumption resolution 280,145] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
282. r3_semi_af1(sK0,sK2,sK1,sK1,sK2) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (2, 3, 4, 5) [subsumption resolution 281,150] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
283. r3_semi_af1(sK0,sK2,sK1,sK1,sK2) | v2_struct_0(sK0) <- (2, 3, 4, 5, 6) [subsumption resolution 282,155] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
284. r3_semi_af1(sK0,sK2,sK1,sK1,sK2) <- (2, 3, 4, 5, 6, ~7) [subsumption resolution 283,160] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
286. 22 <=> r3_semi_af1(sK0,sK2,sK1,sK1,sK2) [avatar definition]
288. r3_semi_af1(sK0,sK2,sK1,sK1,sK2) <- (22) [avatar component clause 286] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
289. 22 | ~2 | ~3 | ~4 | ~5 | ~6 | 7 [avatar split clause 284,158,153,148,143,138,133,286]
292. sK1 = sK2 | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (22) [resolution 288,112] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
294. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~1, 22) [subsumption resolution 292,130] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
295. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~1, 4, 22) [subsumption resolution 294,145] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
296. ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 22) [subsumption resolution 295,140] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
297. ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 5, 22) [subsumption resolution 296,150] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
298. v2_struct_0(sK0) <- (~1, 3, 4, 5, 6, 22) [subsumption resolution 297,155] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
299. $false <- (~1, 3, 4, 5, 6, ~7, 22) [subsumption resolution 298,160] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
300. 1 | ~3 | ~4 | ~5 | ~6 | 7 | ~22 [avatar contradiction clause 299]
301. $false [avatar sat refutation 131,136,141,146,151,156,161,289,300] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:141,thDist:-141}
% SZS output end Proof for semi_af1__t79_semi_af1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 30
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 38
% Generated clauses: 120
% Active clauses: 39
% Passive clauses: 46
% Final active clauses: 39
% Final passive clauses: 2
% Main loop iterations started: 44
% 
% >>> Simplifying Inferences
% Duplicate literals: 29
% Fw subsumption resolutions: 22
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Equational tautologies: 1
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 1
% Forward superposition: 4
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 12
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

