i: [1,49,0,3,1,1,0,0] t123_zfmisc_1
a: [3,77,0,3,1,49]
s: 77
i: [1,51,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,82,0,2,1,51]
s: 82
i: [1,52,0,2,1,0,0,2] fc1_xboole_0
a: [3,87,0,2,1,52]
s: 87
i: [1,67,0,2,1,0,0,2] rc2_xboole_0
a: [3,97,0,2,1,67]
s: 97
i: [1,68,0,2,1,0,0,2] rc1_xboole_0
a: [3,102,0,2,1,68]
s: 102
i: [1,54,0,3,1,0,0,1] fc2_xboole_0
s: 54
i: [1,53,0,3,1,0,0,255] d2_xboole_0
a: [3,92,0,3,1,53]
s: 92
r: 87
i: [1,50,0,15,1,1,0,0] t123_zfmisc_1
a: [3,72,0,15,1,50]
s: 72
i: [1,55,0,5,1,0,0,1] t4_boole
s: 55
i: [1,56,0,5,1,0,0,1] t3_boole
s: 56
i: [1,57,0,5,1,0,0,2] t1_boole
s: 57
i: [1,58,0,5,2,0,0,2] t6_boole
s: 58
i: [1,62,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 62
d: [2,105,0,5,1,1,22,55,92]
s: 105
d: [2,106,0,5,1,1,22,56,92]
s: 106
d: [2,107,0,5,1,1,22,57,92]
s: 107
d: [2,108,0,5,2,1,22,58,92]
s: 108
i: [1,64,0,6,2,0,0,1] fc4_xboole_0
s: 64
i: [1,65,0,6,2,0,0,1] fc5_xboole_0
s: 65
i: [1,66,0,6,2,0,0,2] symmetry_r1_xboole_0
s: 66
i: [1,60,0,7,3,0,0,2] t8_boole
s: 60
i: [1,63,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 63
r: 72
d: [2,118,0,15,1,1,23,72,63]
a: [3,130,0,15,1,118]
s: 130
d: [2,119,1,5,1,1,44,63,107]
s: 119
i: [1,59,0,8,2,0,0,1] t11_zfmisc_1
s: 59
i: [1,61,0,14,2,0,0,2] t87_xboole_1
s: 61
d: [2,146,1,20,2,1,44,130,61]
d: [2,151,1,5,1,1,21,146,63]
a: [3,155,1,5,1,151]
s: 155
d: [2,157,2,0,0,2,52,77,155,59]
e: 157
d: [2,158,2,3,1,1,40,155,59]
d: [2,160,2,0,0,2,21,158,77]
e: 160
f: 160,151,118,49,50
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t123_zfmisc_1
% SZS output start Proof for zfmisc_1__t123_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (X0 != X1 => k4_xboole_0(k2_xboole_0(X2,k1_tarski(X0)),k1_tarski(X1)) = k2_xboole_0(k4_xboole_0(X2,k1_tarski(X1)),k1_tarski(X0))) [input t123_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (X0 != X1 => k4_xboole_0(k2_xboole_0(X2,k1_tarski(X0)),k1_tarski(X1)) = k2_xboole_0(k4_xboole_0(X2,k1_tarski(X1)),k1_tarski(X0))) [negated conjecture 1]
18. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
25. ! [X0] : ! [X1] : (X0 != X1 => r1_xboole_0(k1_tarski(X0),k1_tarski(X1))) [input t11_zfmisc_1]
26. ! [X0] : ! [X1] : ! [X2] : (r1_xboole_0(X0,X1) => k2_xboole_0(k4_xboole_0(X2,X0),X1) = k4_xboole_0(k2_xboole_0(X2,X1),X0)) [input t87_xboole_1]
27. ~! [X0,X1,X2] : (X0 != X1 => k4_xboole_0(k2_xboole_0(X2,k1_tarski(X0)),k1_tarski(X1)) = k2_xboole_0(k4_xboole_0(X2,k1_tarski(X1)),k1_tarski(X0))) [flattening 2]
28. ! [X0,X1] : (X0 != X1 => r1_xboole_0(k1_tarski(X0),k1_tarski(X1))) [flattening 25]
31. ! [X0,X1,X2] : (r1_xboole_0(X0,X1) => k2_xboole_0(k4_xboole_0(X2,X0),X1) = k4_xboole_0(k2_xboole_0(X2,X1),X0)) [flattening 26]
35. ? [X0,X1,X2] : (k4_xboole_0(k2_xboole_0(X2,k1_tarski(X0)),k1_tarski(X1)) != k2_xboole_0(k4_xboole_0(X2,k1_tarski(X1)),k1_tarski(X0)) & X0 != X1) [ennf transformation 27]
37. ! [X0,X1] : (r1_xboole_0(k1_tarski(X0),k1_tarski(X1)) | X0 = X1) [ennf transformation 28]
39. ! [X0,X1,X2] : (k2_xboole_0(k4_xboole_0(X2,X0),X1) = k4_xboole_0(k2_xboole_0(X2,X1),X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 31]
43. ? [X0,X1,X2] : (k4_xboole_0(k2_xboole_0(X2,k1_tarski(X0)),k1_tarski(X1)) != k2_xboole_0(k4_xboole_0(X2,k1_tarski(X1)),k1_tarski(X0)) & X0 != X1) => (k4_xboole_0(k2_xboole_0(sK2,k1_tarski(sK0)),k1_tarski(sK1)) != k2_xboole_0(k4_xboole_0(sK2,k1_tarski(sK1)),k1_tarski(sK0)) & sK0 != sK1) [choice axiom]
44. k4_xboole_0(k2_xboole_0(sK2,k1_tarski(sK0)),k1_tarski(sK1)) != k2_xboole_0(k4_xboole_0(sK2,k1_tarski(sK1)),k1_tarski(sK0)) & sK0 != sK1 [skolemisation 35,43]
49. sK0 != sK1 [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. k4_xboole_0(k2_xboole_0(sK2,k1_tarski(sK0)),k1_tarski(sK1)) != k2_xboole_0(k4_xboole_0(sK2,k1_tarski(sK1)),k1_tarski(sK0)) [cnf transformation 44] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. r1_xboole_0(k1_tarski(X0),k1_tarski(X1)) | X0 = X1 [cnf transformation 37] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. k2_xboole_0(k4_xboole_0(X2,X0),X1) = k4_xboole_0(k2_xboole_0(X2,X1),X0) | ~r1_xboole_0(X0,X1) [cnf transformation 39] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
63. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 18] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. 1 <=> k4_xboole_0(k2_xboole_0(sK2,k1_tarski(sK0)),k1_tarski(sK1)) = k2_xboole_0(k4_xboole_0(sK2,k1_tarski(sK1)),k1_tarski(sK0)) [avatar definition]
72. k4_xboole_0(k2_xboole_0(sK2,k1_tarski(sK0)),k1_tarski(sK1)) != k2_xboole_0(k4_xboole_0(sK2,k1_tarski(sK1)),k1_tarski(sK0)) <- (~1) [avatar component clause 70] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
73. ~1 [avatar split clause 50,70]
75. 2 <=> sK0 = sK1 [avatar definition]
77. sK0 != sK1 <- (~2) [avatar component clause 75] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
78. ~2 [avatar split clause 49,75]
118. k4_xboole_0(k2_xboole_0(sK2,k1_tarski(sK0)),k1_tarski(sK1)) != k2_xboole_0(k1_tarski(sK0),k4_xboole_0(sK2,k1_tarski(sK1))) <- (~1) [backward demodulation 72,63] {a:0,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
128. 8 <=> k4_xboole_0(k2_xboole_0(sK2,k1_tarski(sK0)),k1_tarski(sK1)) = k2_xboole_0(k1_tarski(sK0),k4_xboole_0(sK2,k1_tarski(sK1))) [avatar definition]
130. k4_xboole_0(k2_xboole_0(sK2,k1_tarski(sK0)),k1_tarski(sK1)) != k2_xboole_0(k1_tarski(sK0),k4_xboole_0(sK2,k1_tarski(sK1))) <- (~8) [avatar component clause 128] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
131. ~8 | 1 [avatar split clause 118,70,128]
146. k2_xboole_0(k4_xboole_0(sK2,k1_tarski(sK1)),k1_tarski(sK0)) != k2_xboole_0(k1_tarski(sK0),k4_xboole_0(sK2,k1_tarski(sK1))) | ~r1_xboole_0(k1_tarski(sK1),k1_tarski(sK0)) <- (~8) [superposition 130,61] {a:1,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
151. ~r1_xboole_0(k1_tarski(sK1),k1_tarski(sK0)) <- (~8) [subsumption resolution 146,63] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
153. 9 <=> r1_xboole_0(k1_tarski(sK1),k1_tarski(sK0)) [avatar definition]
155. ~r1_xboole_0(k1_tarski(sK1),k1_tarski(sK0)) <- (~9) [avatar component clause 153] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
156. ~9 | 8 [avatar split clause 151,128,153]
158. sK0 = sK1 <- (~9) [resolution 155,59] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
160. $false <- (~2, ~9) [subsumption resolution 158,77] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
161. 2 | 9 [avatar contradiction clause 160]
162. $false [avatar sat refutation 73,78,131,156,161] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:37,thDist:-37}
% SZS output end Proof for zfmisc_1__t123_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 20
% Generated clauses: 75
% Active clauses: 23
% Passive clauses: 39
% Final active clauses: 21
% Final passive clauses: 12
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 2
% Fw demodulations: 7
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 13
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 15
% Backward superposition: 16
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 7
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

