i: [1,50,0,3,1,1,0,0] t42_zfmisc_1
a: [3,81,0,3,1,50]
s: 81
i: [1,49,0,3,1,1,0,0] t42_zfmisc_1
a: [3,86,0,3,1,49]
s: 86
i: [1,51,0,7,1,1,0,0] t42_zfmisc_1
a: [3,76,0,7,1,51]
s: 76
i: [1,52,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,91,0,2,1,52]
s: 91
i: [1,53,0,2,1,0,0,255] fc1_xboole_0
a: [3,96,0,2,1,53]
s: 96
i: [1,71,0,2,1,0,0,2] rc2_xboole_0
a: [3,106,0,2,1,71]
s: 106
i: [1,72,0,2,1,0,0,2] rc1_xboole_0
a: [3,111,0,2,1,72]
s: 111
i: [1,63,0,3,1,0,0,255] reflexivity_r1_tarski
s: 63
i: [1,54,0,3,1,0,0,255] d2_xboole_0
a: [3,101,0,3,1,54]
s: 101
r: 96
i: [1,64,0,4,1,0,0,1] fc3_xboole_0
s: 64
i: [1,55,0,5,1,0,0,255] t1_boole
s: 55
i: [1,56,0,5,2,0,0,255] t6_boole
s: 56
i: [1,59,0,5,2,0,0,1] t7_boole
s: 59
d: [2,116,1,2,1,1,40,59,81]
a: [3,121,1,2,1,116]
s: 121
i: [1,65,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 65
d: [2,114,0,5,1,1,22,55,101]
s: 114
d: [2,115,0,5,2,1,22,56,101]
s: 115
i: [1,68,0,6,2,0,0,1] fc5_xboole_0
s: 68
i: [1,69,0,6,2,0,0,1] fc4_xboole_0
s: 69
i: [1,70,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 70
i: [1,58,0,7,3,0,0,1] t8_boole
s: 58
d: [2,131,1,3,1,1,40,70,81]
a: [3,136,1,3,1,131]
s: 136
i: [1,66,0,7,1,0,0,1] commutativity_k2_tarski
s: 66
d: [2,132,1,3,1,1,40,70,86]
a: [3,141,1,3,1,132]
s: 141
i: [1,67,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 67
i: [1,57,0,8,2,0,0,255] t12_xboole_1
s: 57
d: [2,164,1,8,2,1,44,76,57]
d: [2,170,1,5,1,1,19,164]
a: [3,174,1,5,1,170]
s: 174
i: [1,60,0,8,2,0,0,255] t32_zfmisc_1
s: 60
d: [2,147,1,7,1,1,44,76,67]
a: [3,158,1,7,1,147]
s: 158
i: [1,61,0,8,2,0,0,255] t32_zfmisc_1
s: 61
i: [1,62,0,11,3,0,0,255] t32_zfmisc_1
s: 62
d: [2,192,2,0,0,3,52,86,81,174,62]
e: 192
d: [2,193,2,6,2,1,40,62,174]
d: [2,199,2,3,1,2,21,193,81]
d: [2,200,2,0,0,3,21,199,86]
e: 200
f: 200,170,49,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t42_zfmisc_1
% SZS output start Proof for zfmisc_1__t42_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ((r2_hidden(X2,X1) & r2_hidden(X0,X1)) => k2_xboole_0(k2_tarski(X0,X2),X1) = X1) [input t42_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ((r2_hidden(X2,X1) & r2_hidden(X0,X1)) => k2_xboole_0(k2_tarski(X0,X2),X1) = X1) [negated conjecture 1]
23. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k2_xboole_0(X0,X1) = X1) [input t12_xboole_1]
24. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),X2) <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [input t32_zfmisc_1]
25. ~! [X0,X1,X2] : ((r2_hidden(X2,X1) & r2_hidden(X0,X1)) => k2_xboole_0(k2_tarski(X0,X2),X1) = X1) [flattening 2]
26. ! [X0,X1] : (r1_tarski(X0,X1) => k2_xboole_0(X0,X1) = X1) [flattening 23]
29. ! [X0,X1,X2] : (r1_tarski(k2_tarski(X0,X1),X2) <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [flattening 24]
32. ? [X0,X1,X2] : (k2_xboole_0(k2_tarski(X0,X2),X1) != X1 & (r2_hidden(X2,X1) & r2_hidden(X0,X1))) [ennf transformation 25]
33. ? [X0,X1,X2] : (k2_xboole_0(k2_tarski(X0,X2),X1) != X1 & r2_hidden(X2,X1) & r2_hidden(X0,X1)) [flattening 32]
35. ! [X0,X1] : (k2_xboole_0(X0,X1) = X1 | ~r1_tarski(X0,X1)) [ennf transformation 26]
41. ? [X0,X1,X2] : (k2_xboole_0(k2_tarski(X0,X2),X1) != X1 & r2_hidden(X2,X1) & r2_hidden(X0,X1)) => (sK1 != k2_xboole_0(k2_tarski(sK0,sK2),sK1) & r2_hidden(sK2,sK1) & r2_hidden(sK0,sK1)) [choice axiom]
42. sK1 != k2_xboole_0(k2_tarski(sK0,sK2),sK1) & r2_hidden(sK2,sK1) & r2_hidden(sK0,sK1) [skolemisation 33,41]
43. ! [X0,X1,X2] : ((r1_tarski(k2_tarski(X0,X1),X2) | (~r2_hidden(X1,X2) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | ~r1_tarski(k2_tarski(X0,X1),X2))) [nnf transformation 29]
44. ! [X0,X1,X2] : ((r1_tarski(k2_tarski(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | ~r1_tarski(k2_tarski(X0,X1),X2))) [flattening 43]
49. r2_hidden(sK0,sK1) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. r2_hidden(sK2,sK1) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. sK1 != k2_xboole_0(k2_tarski(sK0,sK2),sK1) [cnf transformation 42] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
57. k2_xboole_0(X0,X1) = X1 | ~r1_tarski(X0,X1) [cnf transformation 35] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
62. r1_tarski(k2_tarski(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) [cnf transformation 44] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
74. 1 <=> sK1 = k2_xboole_0(k2_tarski(sK0,sK2),sK1) [avatar definition]
76. sK1 != k2_xboole_0(k2_tarski(sK0,sK2),sK1) <- (~1) [avatar component clause 74] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
77. ~1 [avatar split clause 51,74]
79. 2 <=> r2_hidden(sK2,sK1) [avatar definition]
81. r2_hidden(sK2,sK1) <- (2) [avatar component clause 79] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
82. 2 [avatar split clause 50,79]
84. 3 <=> r2_hidden(sK0,sK1) [avatar definition]
86. r2_hidden(sK0,sK1) <- (3) [avatar component clause 84] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. 3 [avatar split clause 49,84]
164. sK1 != sK1 | ~r1_tarski(k2_tarski(sK0,sK2),sK1) <- (~1) [superposition 76,57] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
170. ~r1_tarski(k2_tarski(sK0,sK2),sK1) <- (~1) [trivial inequality removal 164] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
172. 13 <=> r1_tarski(k2_tarski(sK0,sK2),sK1) [avatar definition]
174. ~r1_tarski(k2_tarski(sK0,sK2),sK1) <- (~13) [avatar component clause 172] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
175. ~13 | 1 [avatar split clause 170,74,172]
193. ~r2_hidden(sK2,sK1) | ~r2_hidden(sK0,sK1) <- (~13) [resolution 62,174] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
199. ~r2_hidden(sK0,sK1) <- (2, ~13) [subsumption resolution 193,81] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
200. $false <- (2, 3, ~13) [subsumption resolution 199,86] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
201. ~2 | ~3 | 13 [avatar contradiction clause 200]
202. $false [avatar sat refutation 77,82,87,175,201] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:75,thDist:-75}
% SZS output end Proof for zfmisc_1__t42_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 89
% Active clauses: 29
% Passive clauses: 49
% Final active clauses: 28
% Final passive clauses: 18
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 2
% Fw demodulations: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 1
% Forward superposition: 16
% Backward superposition: 18
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 8
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% passive container maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

