i: [1,39,0,2,1,1,0,0] t1_projdes1
a: [3,97,0,2,1,39]
s: 97
i: [1,38,0,2,1,1,0,0] t1_projdes1
a: [3,102,0,2,1,38]
s: 102
i: [1,37,0,2,1,1,0,0] t1_projdes1
a: [3,107,0,2,1,37]
s: 107
i: [1,36,0,2,1,1,0,0] t1_projdes1
a: [3,112,0,2,1,36]
s: 112
i: [1,42,0,4,1,1,0,0] t1_projdes1
a: [3,82,0,4,1,42]
s: 82
i: [1,41,0,4,1,1,0,0] t1_projdes1
a: [3,87,0,4,1,41]
s: 87
i: [1,40,0,4,1,1,0,0] t1_projdes1
a: [3,92,0,4,1,40]
s: 92
i: [1,44,0,25,5,1,0,0] t1_projdes1
a: [3,56,0,5,1,44]
s: 56
i: [1,43,0,5,1,1,0,0] t1_projdes1
a: [3,77,0,5,1,43]
s: 77
i: [1,48,0,2,1,0,0,1] existence_l1_collsp
a: [3,117,0,2,1,48]
s: 117
i: [1,52,0,2,1,0,0,1] rc2_anproj_2
a: [3,122,0,2,1,52]
s: 122
i: [1,51,0,2,1,0,0,1] rc2_anproj_2
a: [3,127,0,2,1,51]
s: 127
i: [1,50,0,2,1,0,0,1] rc2_anproj_2
a: [3,132,0,2,1,50]
s: 132
i: [1,49,0,2,1,0,0,1] rc2_anproj_2
a: [3,137,0,2,1,49]
s: 137
i: [1,47,0,4,1,0,0,1] existence_m1_subset_1
s: 47
i: [1,45,0,30,9,0,0,1] t4_collsp
s: 45
a: [3,63,0,5,1,44]
s: 63
i: [1,46,0,30,9,0,0,1] t4_collsp
s: 46
d: [2,149,1,0,0,9,52,112,107,102,97,87,92,63,56,82,46]
e: 149
d: [2,151,1,25,8,1,40,46,63]
d: [2,160,1,20,7,2,21,151,56]
d: [2,161,1,16,6,3,21,160,82]
d: [2,162,1,12,5,4,21,161,92]
d: [2,163,1,8,4,5,21,162,87]
d: [2,164,1,6,3,6,21,163,97]
d: [2,165,1,4,2,7,21,164,102]
d: [2,166,1,2,1,8,21,165,107]
d: [2,167,1,0,0,9,21,166,112]
e: 167
r: 56
a: [3,55,0,5,1,44]
s: 55
a: [3,60,0,5,1,44]
s: 60
a: [3,67,0,5,1,44]
s: 67
d: [2,188,1,0,0,9,52,112,107,102,97,92,82,87,60,67,45]
e: 188
d: [2,190,1,25,8,1,40,67,45]
d: [2,192,1,20,7,2,21,190,60]
d: [2,193,1,16,6,3,21,192,87]
d: [2,194,1,12,5,4,21,193,82]
d: [2,195,1,8,4,5,21,194,92]
d: [2,196,1,6,3,6,21,195,97]
d: [2,197,1,4,2,7,21,196,102]
d: [2,198,1,2,1,8,21,197,107]
d: [2,199,1,0,0,9,21,198,112]
e: 199
d: [2,139,1,25,8,1,40,45,77]
d: [2,140,1,21,7,2,21,139,82]
d: [2,141,1,17,6,3,21,140,87]
d: [2,142,1,13,5,4,21,141,92]
d: [2,143,1,11,4,5,21,142,97]
d: [2,144,1,9,3,6,21,143,102]
d: [2,145,1,7,2,7,21,144,107]
d: [2,146,1,5,1,8,21,145,112]
d: [2,169,1,21,7,2,21,151,82]
d: [2,170,1,17,6,3,21,169,92]
d: [2,171,1,13,5,4,21,170,87]
d: [2,172,1,11,4,5,21,171,97]
d: [2,173,1,9,3,6,21,172,102]
d: [2,174,1,7,2,7,21,173,107]
d: [2,175,1,5,1,8,21,174,112]
d: [2,150,1,25,8,1,40,46,77]
d: [2,153,1,21,7,2,21,150,82]
d: [2,154,1,17,6,3,21,153,87]
d: [2,155,1,13,5,4,21,154,92]
d: [2,156,1,11,4,5,21,155,97]
d: [2,157,1,9,3,6,21,156,102]
d: [2,158,1,7,2,7,21,157,107]
d: [2,159,1,5,1,8,21,158,112]
d: [2,179,1,25,8,1,40,55,45]
d: [2,180,1,21,7,2,21,179,92]
d: [2,181,1,17,6,3,21,180,82]
d: [2,182,1,13,5,4,21,181,87]
d: [2,183,1,11,4,5,21,182,97]
d: [2,184,1,9,3,6,21,183,102]
d: [2,185,1,7,2,7,21,184,107]
d: [2,186,1,5,1,8,21,185,112]
f: 199,186,159,175,146,36,37,38,39,40,41,42,43,44
% Refutation found. Thanks to Tanya!
% SZS status Theorem for projdes1__t1_projdes1
% SZS output start Proof for projdes1__t1_projdes1
1. ! [X0] : ((l1_collsp(X0) & ~v7_anproj_2(X0) & v3_anproj_2(X0) & v2_anproj_2(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [input t1_projdes1]
2. ~! [X0] : ((l1_collsp(X0) & ~v7_anproj_2(X0) & v3_anproj_2(X0) & v2_anproj_2(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [negated conjecture 1]
11. ! [X0] : ((l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3))))))) [input t4_collsp]
14. ~! [X0] : ((l1_collsp(X0) & v3_anproj_2(X0) & v2_anproj_2(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [pure predicate removal 2]
15. ~! [X0] : ((l1_collsp(X0) & v2_anproj_2(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [pure predicate removal 14]
17. ~! [X0] : ((l1_collsp(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [pure predicate removal 15]
19. ~! [X0] : ((l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [pure predicate removal 17]
21. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~r1_collsp(X0,X3,X2,X1) | ~r1_collsp(X0,X1,X3,X2) | ~r1_collsp(X0,X2,X1,X3) | ~r1_collsp(X0,X3,X1,X2) | ~r1_collsp(X0,X2,X3,X1)) & r1_collsp(X0,X1,X2,X3)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0))) [ennf transformation 19]
22. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_collsp(X0,X3,X2,X1) | ~r1_collsp(X0,X1,X3,X2) | ~r1_collsp(X0,X2,X1,X3) | ~r1_collsp(X0,X3,X1,X2) | ~r1_collsp(X0,X2,X3,X1)) & r1_collsp(X0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) [flattening 21]
23. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3)) | ~r1_collsp(X0,X1,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0))) [ennf transformation 11]
24. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3)) | ~r1_collsp(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0)) [flattening 23]
25. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_collsp(X0,X3,X2,X1) | ~r1_collsp(X0,X1,X3,X2) | ~r1_collsp(X0,X2,X1,X3) | ~r1_collsp(X0,X3,X1,X2) | ~r1_collsp(X0,X2,X3,X1)) & r1_collsp(X0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : ((~r1_collsp(sK0,X3,X2,X1) | ~r1_collsp(sK0,X1,X3,X2) | ~r1_collsp(sK0,X2,X1,X3) | ~r1_collsp(sK0,X3,X1,X2) | ~r1_collsp(sK0,X2,X3,X1)) & r1_collsp(sK0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_collsp(sK0) & v3_collsp(sK0) & v2_collsp(sK0) & ~v2_struct_0(sK0)) [choice axiom]
26. ? [X1] : (? [X2] : (? [X3] : ((~r1_collsp(sK0,X3,X2,X1) | ~r1_collsp(sK0,X1,X3,X2) | ~r1_collsp(sK0,X2,X1,X3) | ~r1_collsp(sK0,X3,X1,X2) | ~r1_collsp(sK0,X2,X3,X1)) & r1_collsp(sK0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : ((~r1_collsp(sK0,X3,X2,sK1) | ~r1_collsp(sK0,sK1,X3,X2) | ~r1_collsp(sK0,X2,sK1,X3) | ~r1_collsp(sK0,X3,sK1,X2) | ~r1_collsp(sK0,X2,X3,sK1)) & r1_collsp(sK0,sK1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
27. ? [X2] : (? [X3] : ((~r1_collsp(sK0,X3,X2,sK1) | ~r1_collsp(sK0,sK1,X3,X2) | ~r1_collsp(sK0,X2,sK1,X3) | ~r1_collsp(sK0,X3,sK1,X2) | ~r1_collsp(sK0,X2,X3,sK1)) & r1_collsp(sK0,sK1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : ((~r1_collsp(sK0,X3,sK2,sK1) | ~r1_collsp(sK0,sK1,X3,sK2) | ~r1_collsp(sK0,sK2,sK1,X3) | ~r1_collsp(sK0,X3,sK1,sK2) | ~r1_collsp(sK0,sK2,X3,sK1)) & r1_collsp(sK0,sK1,sK2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
28. ? [X3] : ((~r1_collsp(sK0,X3,sK2,sK1) | ~r1_collsp(sK0,sK1,X3,sK2) | ~r1_collsp(sK0,sK2,sK1,X3) | ~r1_collsp(sK0,X3,sK1,sK2) | ~r1_collsp(sK0,sK2,X3,sK1)) & r1_collsp(sK0,sK1,sK2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) => ((~r1_collsp(sK0,sK3,sK2,sK1) | ~r1_collsp(sK0,sK1,sK3,sK2) | ~r1_collsp(sK0,sK2,sK1,sK3) | ~r1_collsp(sK0,sK3,sK1,sK2) | ~r1_collsp(sK0,sK2,sK3,sK1)) & r1_collsp(sK0,sK1,sK2,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
29. ((((~r1_collsp(sK0,sK3,sK2,sK1) | ~r1_collsp(sK0,sK1,sK3,sK2) | ~r1_collsp(sK0,sK2,sK1,sK3) | ~r1_collsp(sK0,sK3,sK1,sK2) | ~r1_collsp(sK0,sK2,sK3,sK1)) & r1_collsp(sK0,sK1,sK2,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_collsp(sK0) & v3_collsp(sK0) & v2_collsp(sK0) & ~v2_struct_0(sK0) [skolemisation 22,28,27,26,25]
36. ~v2_struct_0(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
37. v2_collsp(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
38. v3_collsp(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
39. l1_collsp(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
40. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
41. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
42. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
43. r1_collsp(sK0,sK1,sK2,sK3) [cnf transformation 29] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
44. ~r1_collsp(sK0,sK3,sK2,sK1) | ~r1_collsp(sK0,sK1,sK3,sK2) | ~r1_collsp(sK0,sK2,sK1,sK3) | ~r1_collsp(sK0,sK3,sK1,sK2) | ~r1_collsp(sK0,sK2,sK3,sK1) [cnf transformation 29] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
45. ~r1_collsp(X0,X1,X2,X3) | r1_collsp(X0,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0) [cnf transformation 24] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
46. ~r1_collsp(X0,X1,X2,X3) | r1_collsp(X0,X1,X3,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0) [cnf transformation 24] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. 1 <=> r1_collsp(sK0,sK2,sK3,sK1) [avatar definition]
55. r1_collsp(sK0,sK2,sK3,sK1) <- (1) [avatar component clause 54] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
58. 2 <=> r1_collsp(sK0,sK3,sK1,sK2) [avatar definition]
60. ~r1_collsp(sK0,sK3,sK1,sK2) <- (~2) [avatar component clause 58] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
62. 3 <=> r1_collsp(sK0,sK2,sK1,sK3) [avatar definition]
63. r1_collsp(sK0,sK2,sK1,sK3) <- (3) [avatar component clause 62] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
66. 4 <=> r1_collsp(sK0,sK1,sK3,sK2) [avatar definition]
67. r1_collsp(sK0,sK1,sK3,sK2) <- (4) [avatar component clause 66] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
70. 5 <=> r1_collsp(sK0,sK3,sK2,sK1) [avatar definition]
73. ~1 | ~2 | ~3 | ~4 | ~5 [avatar split clause 44,70,66,62,58,54]
75. 6 <=> r1_collsp(sK0,sK1,sK2,sK3) [avatar definition]
77. r1_collsp(sK0,sK1,sK2,sK3) <- (6) [avatar component clause 75] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
78. 6 [avatar split clause 43,75]
80. 7 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
82. m1_subset_1(sK3,u1_struct_0(sK0)) <- (7) [avatar component clause 80] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
83. 7 [avatar split clause 42,80]
85. 8 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
87. m1_subset_1(sK2,u1_struct_0(sK0)) <- (8) [avatar component clause 85] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
88. 8 [avatar split clause 41,85]
90. 9 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
92. m1_subset_1(sK1,u1_struct_0(sK0)) <- (9) [avatar component clause 90] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
93. 9 [avatar split clause 40,90]
95. 10 <=> l1_collsp(sK0) [avatar definition]
97. l1_collsp(sK0) <- (10) [avatar component clause 95] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
98. 10 [avatar split clause 39,95]
100. 11 <=> v3_collsp(sK0) [avatar definition]
102. v3_collsp(sK0) <- (11) [avatar component clause 100] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
103. 11 [avatar split clause 38,100]
105. 12 <=> v2_collsp(sK0) [avatar definition]
107. v2_collsp(sK0) <- (12) [avatar component clause 105] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
108. 12 [avatar split clause 37,105]
110. 13 <=> v2_struct_0(sK0) [avatar definition]
112. ~v2_struct_0(sK0) <- (~13) [avatar component clause 110] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
113. ~13 [avatar split clause 36,110]
139. r1_collsp(sK0,sK2,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6) [resolution 45,77] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
140. r1_collsp(sK0,sK2,sK1,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7) [subsumption resolution 139,82] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
141. r1_collsp(sK0,sK2,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8) [subsumption resolution 140,87] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
142. r1_collsp(sK0,sK2,sK1,sK3) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9) [subsumption resolution 141,92] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
143. r1_collsp(sK0,sK2,sK1,sK3) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10) [subsumption resolution 142,97] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
144. r1_collsp(sK0,sK2,sK1,sK3) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10, 11) [subsumption resolution 143,102] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
145. r1_collsp(sK0,sK2,sK1,sK3) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10, 11, 12) [subsumption resolution 144,107] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
146. r1_collsp(sK0,sK2,sK1,sK3) <- (6, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 145,112] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
147. 3 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar split clause 146,110,105,100,95,90,85,80,75,62]
150. r1_collsp(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6) [resolution 46,77] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
151. r1_collsp(sK0,sK2,sK3,sK1) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3) [resolution 46,63] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
153. r1_collsp(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7) [subsumption resolution 150,82] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
154. r1_collsp(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8) [subsumption resolution 153,87] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
155. r1_collsp(sK0,sK1,sK3,sK2) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9) [subsumption resolution 154,92] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
156. r1_collsp(sK0,sK1,sK3,sK2) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10) [subsumption resolution 155,97] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
157. r1_collsp(sK0,sK1,sK3,sK2) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10, 11) [subsumption resolution 156,102] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
158. r1_collsp(sK0,sK1,sK3,sK2) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10, 11, 12) [subsumption resolution 157,107] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
159. r1_collsp(sK0,sK1,sK3,sK2) <- (6, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 158,112] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
169. r1_collsp(sK0,sK2,sK3,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7) [subsumption resolution 151,82] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
170. r1_collsp(sK0,sK2,sK3,sK1) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7, 9) [subsumption resolution 169,92] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
171. r1_collsp(sK0,sK2,sK3,sK1) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7, 8, 9) [subsumption resolution 170,87] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
172. r1_collsp(sK0,sK2,sK3,sK1) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7, 8, 9, 10) [subsumption resolution 171,97] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
173. r1_collsp(sK0,sK2,sK3,sK1) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7, 8, 9, 10, 11) [subsumption resolution 172,102] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
174. r1_collsp(sK0,sK2,sK3,sK1) | v2_struct_0(sK0) <- (3, 7, 8, 9, 10, 11, 12) [subsumption resolution 173,107] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
175. r1_collsp(sK0,sK2,sK3,sK1) <- (3, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 174,112] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
176. 1 | ~3 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar split clause 175,110,105,100,95,90,85,80,62,54]
177. 4 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar split clause 159,110,105,100,95,90,85,80,75,66]
179. r1_collsp(sK0,sK3,sK2,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1) [resolution 55,45] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
180. r1_collsp(sK0,sK3,sK2,sK1) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 9) [subsumption resolution 179,92] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
181. r1_collsp(sK0,sK3,sK2,sK1) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 7, 9) [subsumption resolution 180,82] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
182. r1_collsp(sK0,sK3,sK2,sK1) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 7, 8, 9) [subsumption resolution 181,87] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
183. r1_collsp(sK0,sK3,sK2,sK1) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 7, 8, 9, 10) [subsumption resolution 182,97] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
184. r1_collsp(sK0,sK3,sK2,sK1) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 7, 8, 9, 10, 11) [subsumption resolution 183,102] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
185. r1_collsp(sK0,sK3,sK2,sK1) | v2_struct_0(sK0) <- (1, 7, 8, 9, 10, 11, 12) [subsumption resolution 184,107] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
186. r1_collsp(sK0,sK3,sK2,sK1) <- (1, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 185,112] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
187. 5 | ~1 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar split clause 186,110,105,100,95,90,85,80,54,70]
190. r1_collsp(sK0,sK3,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (4) [resolution 67,45] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
192. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4) [subsumption resolution 190,60] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
193. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 8) [subsumption resolution 192,87] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
194. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 7, 8) [subsumption resolution 193,82] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
195. ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 7, 8, 9) [subsumption resolution 194,92] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
196. ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 7, 8, 9, 10) [subsumption resolution 195,97] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
197. ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 7, 8, 9, 10, 11) [subsumption resolution 196,102] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
198. v2_struct_0(sK0) <- (~2, 4, 7, 8, 9, 10, 11, 12) [subsumption resolution 197,107] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
199. $false <- (~2, 4, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 198,112] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
200. 2 | ~4 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar contradiction clause 199]
201. $false [avatar sat refutation 73,78,83,88,93,98,103,108,113,147,176,177,187,200] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:237,thDist:-237}
% SZS output end Proof for projdes1__t1_projdes1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 5
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 92
% Active clauses: 21
% Passive clauses: 22
% Final active clauses: 20
% Final passive clauses: 1
% Main loop iterations started: 21
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 44
% 
% >>> Deletion Inferences
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 8
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 5
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 13
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% backward subsumption index maintenance: 0.001 s ( own 0.0000 s ) 
% other: 0.002 s ( own 0.001 s ) 

