i: [1,245,0,2,1,1,0,0] t29_sin_cos7
a: [3,388,0,2,1,245]
s: 388
i: [1,246,0,3,1,1,0,0] t29_sin_cos7
a: [3,383,0,3,1,246]
s: 383
i: [1,247,0,4,1,1,0,0] t29_sin_cos7
a: [3,378,0,4,1,247]
s: 378
i: [1,248,0,2,1,0,0,3] spc2_boole
a: [3,393,0,2,1,248]
s: 393
i: [1,249,0,2,1,0,0,1] spc1_boole
a: [3,398,0,2,1,249]
s: 398
i: [1,250,0,2,1,0,0,255] fc3_membered
a: [3,403,0,2,1,250]
s: 403
i: [1,251,0,2,1,0,0,3] fc6_membered
a: [3,408,0,2,1,251]
s: 408
i: [1,252,0,2,1,0,0,255] dt_k7_power
a: [3,413,0,2,1,252]
s: 413
i: [1,253,0,2,1,0,0,2] fc1_xboole_0
a: [3,418,0,2,1,253]
s: 418
i: [1,254,0,2,1,0,0,3] spc0_boole
a: [3,423,0,2,1,254]
s: 423
i: [1,273,0,2,1,0,0,2] spc2_numerals
a: [3,533,0,2,1,273]
s: 533
i: [1,277,0,2,1,0,0,2] spc1_numerals
a: [3,553,0,2,1,277]
s: 553
i: [1,350,0,2,1,0,0,2] rc2_xboole_0
a: [3,580,0,2,1,350]
s: 580
i: [1,351,0,2,1,0,0,1] rc1_xreal_0
a: [3,585,0,2,1,351]
s: 585
i: [1,352,0,2,1,0,0,4] rc1_xxreal_0
a: [3,590,0,2,1,352]
s: 590
i: [1,353,0,2,1,0,0,2] rc1_xboole_0
a: [3,595,0,2,1,353]
s: 595
i: [1,355,0,2,1,0,0,3] rc1_membered
a: [3,600,0,2,1,355]
s: 600
i: [1,354,0,2,1,0,0,3] rc1_membered
a: [3,605,0,2,1,354]
s: 605
i: [1,357,0,2,1,0,0,3] rc2_membered
a: [3,610,0,2,1,357]
s: 610
i: [1,356,0,2,1,0,0,3] rc2_membered
a: [3,615,0,2,1,356]
s: 615
i: [1,359,0,2,1,0,0,4] rc3_membered
a: [3,620,0,2,1,359]
s: 620
i: [1,358,0,2,1,0,0,4] rc3_membered
a: [3,625,0,2,1,358]
s: 625
i: [1,362,0,2,1,0,0,255] rc3_xreal_0
a: [3,630,0,2,1,362]
s: 630
i: [1,361,0,2,1,0,0,255] rc3_xreal_0
a: [3,635,0,2,1,361]
s: 635
i: [1,360,0,2,1,0,0,255] rc3_xreal_0
a: [3,640,0,2,1,360]
s: 640
i: [1,365,0,2,1,0,0,255] rc2_xreal_0
a: [3,645,0,2,1,365]
s: 645
i: [1,364,0,2,1,0,0,255] rc2_xreal_0
a: [3,650,0,2,1,364]
s: 650
i: [1,363,0,2,1,0,0,255] rc2_xreal_0
a: [3,655,0,2,1,363]
s: 655
i: [1,367,0,2,1,0,0,255] rc3_xxreal_0
a: [3,660,0,2,1,367]
s: 660
i: [1,366,0,2,1,0,0,255] rc3_xxreal_0
a: [3,665,0,2,1,366]
s: 665
i: [1,369,0,2,1,0,0,3] rc2_xxreal_0
a: [3,670,0,2,1,369]
s: 670
i: [1,368,0,2,1,0,0,3] rc2_xxreal_0
a: [3,675,0,2,1,368]
s: 675
i: [1,372,0,2,1,0,0,255] rc4_xreal_0
a: [3,680,0,2,1,372]
s: 680
i: [1,371,0,2,1,0,0,255] rc4_xreal_0
a: [3,685,0,2,1,371]
s: 685
i: [1,370,0,2,1,0,0,255] rc4_xreal_0
a: [3,690,0,2,1,370]
s: 690
i: [1,374,0,2,1,0,0,4] rc4_xxreal_0
a: [3,695,0,2,1,374]
s: 695
i: [1,373,0,2,1,0,0,4] rc4_xxreal_0
a: [3,700,0,2,1,373]
s: 700
i: [1,341,0,3,1,0,0,255] reflexivity_r1_tarski
s: 341
i: [1,255,0,3,1,0,0,2] l1_sin_cos7
a: [3,428,0,3,1,255]
s: 428
i: [1,256,0,3,1,0,0,2] t11_taylor_1
a: [3,433,0,3,1,256]
s: 433
i: [1,257,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r2_r0
a: [3,438,0,3,1,257]
s: 438
i: [1,258,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r2_r1
a: [3,443,0,3,1,258]
s: 443
i: [1,259,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r1_r0
a: [3,448,0,3,1,259]
s: 448
i: [1,260,0,3,1,0,0,255] redefinition_k8_power
a: [3,453,0,3,1,260]
s: 453
r: 433
i: [1,261,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,458,0,3,1,261]
s: 458
r: 418
d: [2,714,0,2,1,2,23,418,458]
a: [3,718,0,2,1,714]
s: 718
i: [1,262,0,3,1,0,0,3] redefinition_k5_numbers
a: [3,463,0,3,1,262]
s: 463
r: 408
d: [2,720,0,2,1,2,23,408,463]
a: [3,724,0,2,1,720]
s: 724
i: [1,263,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r0_r0
a: [3,468,0,3,1,263]
s: 468
i: [1,264,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r0_r2
a: [3,473,0,3,1,264]
s: 473
i: [1,265,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r0_r1
a: [3,478,0,3,1,265]
s: 478
i: [1,266,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r2_r2
a: [3,483,0,3,1,266]
s: 483
i: [1,267,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r1_r2
a: [3,488,0,3,1,267]
s: 488
i: [1,268,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,493,0,3,1,268]
s: 493
i: [1,269,0,3,1,0,0,2] t1_numerals
a: [3,498,0,3,1,269]
s: 498
i: [1,270,0,3,1,0,0,2] dt_k8_power
a: [3,503,0,3,1,270]
s: 503
i: [1,276,0,3,1,0,0,2] spc2_numerals
a: [3,518,0,3,1,276]
s: 518
i: [1,275,0,3,1,0,0,2] spc2_numerals
a: [3,523,0,3,1,275]
s: 523
i: [1,280,0,3,1,0,0,2] spc1_numerals
a: [3,538,0,3,1,280]
s: 538
i: [1,279,0,3,1,0,0,2] spc1_numerals
a: [3,543,0,3,1,279]
s: 543
i: [1,283,0,3,1,0,0,2] spc0_numerals
a: [3,558,0,3,1,283]
s: 558
i: [1,282,0,3,1,0,0,2] spc0_numerals
a: [3,563,0,3,1,282]
s: 563
d: [2,702,0,3,1,2,23,433,453]
a: [3,707,0,3,1,702]
s: 707
d: [2,726,0,3,1,2,22,498,463]
d: [2,727,0,3,1,3,22,726,458]
a: [3,731,0,3,1,727]
s: 731
d: [2,733,0,3,1,2,22,503,453]
a: [3,737,0,3,1,733]
s: 737
i: [1,284,0,4,2,0,0,255] cc2_xreal_0
s: 284
d: [2,703,1,3,1,2,44,428,453]
a: [3,712,1,3,1,703]
s: 712
i: [1,285,0,4,2,0,0,255] cc2_xxreal_0
s: 285
i: [1,289,0,4,2,0,0,255] cc2_membered
s: 289
i: [1,290,0,4,2,0,0,255] cc3_membered
s: 290
i: [1,293,0,4,2,0,0,255] cc5_membered
s: 293
d: [2,739,1,2,1,1,40,293,403]
a: [3,743,1,2,1,739]
s: 743
i: [1,294,0,4,2,0,0,255] cc4_membered
s: 294
d: [2,745,1,2,1,1,40,294,403]
a: [3,749,1,2,1,745]
s: 749
i: [1,297,0,4,2,0,0,255] cc1_membered
s: 297
i: [1,300,0,4,2,0,0,4] cc4_xreal_0
s: 300
d: [2,776,1,2,1,1,40,300,388]
a: [3,789,1,2,1,776]
s: 789
i: [1,312,0,4,2,0,0,3] cc18_membered
s: 312
d: [2,751,1,2,1,1,40,297,724]
a: [3,758,1,2,1,751]
s: 758
i: [1,340,0,4,1,0,0,3] existence_m1_subset_1
s: 340
i: [1,271,0,4,1,0,0,2] dt_k5_numbers
a: [3,508,0,4,1,271]
s: 508
d: [2,752,1,2,1,1,40,297,600]
a: [3,763,1,2,1,752]
s: 763
i: [1,272,0,4,1,0,0,1] dt_k6_numbers
a: [3,513,0,4,1,272]
s: 513
d: [2,753,1,2,1,1,40,297,610]
a: [3,768,1,2,1,753]
s: 768
i: [1,274,0,4,1,0,0,2] spc2_numerals
a: [3,528,0,4,1,274]
s: 528
i: [1,278,0,4,1,0,0,2] spc1_numerals
a: [3,548,0,4,1,278]
s: 548
d: [2,754,1,2,1,1,40,297,620]
a: [3,773,1,2,1,754]
s: 773
i: [1,281,0,4,1,0,0,2] spc0_numerals
a: [3,568,0,4,1,281]
s: 568
d: [2,775,1,2,1,1,40,300,413]
a: [3,784,1,2,1,775]
s: 784
i: [1,301,0,5,2,0,0,1] fc5_sin_cos
s: 301
i: [1,313,0,5,2,0,0,2] t6_boole
s: 313
d: [2,851,1,4,2,1,44,301,378]
d: [2,852,1,2,1,2,21,851,388]
a: [3,856,1,2,1,852]
s: 856
i: [1,314,0,5,2,0,0,2] cc25_membered
s: 314
d: [2,859,2,2,1,1,40,856,300]
a: [3,863,2,2,1,859]
s: 863
i: [1,315,0,5,2,0,0,2] cc1_xreal_0
s: 315
i: [1,337,0,5,2,0,0,255] t7_boole
s: 337
d: [2,777,1,2,1,1,40,300,585]
a: [3,794,1,2,1,777]
s: 794
i: [1,348,0,7,3,0,0,4] reflexivity_r1_xxreal_0
a: [3,575,0,5,2,348]
s: 575
d: [2,796,1,2,1,1,40,312,718]
a: [3,804,1,2,1,796]
s: 804
d: [2,858,0,5,2,1,22,313,458]
s: 858
i: [1,316,0,6,2,0,0,2] cc11_membered
s: 316
d: [2,797,1,2,1,1,40,312,423]
a: [3,809,1,2,1,797]
s: 809
i: [1,317,0,6,2,0,0,255] cc8_membered
s: 317
d: [2,798,1,2,1,1,40,312,595]
a: [3,814,1,2,1,798]
s: 814
i: [1,324,0,6,3,0,0,255] cc5_xxreal_0
s: 324
i: [1,326,0,6,3,0,0,255] cc5_xxreal_0
s: 326
d: [2,799,1,2,1,1,40,312,690]
a: [3,819,1,2,1,799]
s: 819
i: [1,327,0,6,3,0,0,255] cc3_xxreal_0
s: 327
d: [2,800,1,2,1,1,40,312,700]
a: [3,824,1,2,1,800]
s: 824
i: [1,329,0,6,3,0,0,255] cc3_xxreal_0
s: 329
i: [1,331,0,6,3,0,0,255] cc7_xxreal_0
s: 331
d: [2,866,1,2,1,1,40,314,543]
a: [3,874,1,2,1,866]
s: 874
i: [1,332,0,6,3,0,0,255] cc7_xxreal_0
s: 332
d: [2,868,1,2,1,1,40,314,523]
a: [3,880,1,2,1,868]
s: 880
i: [1,333,0,6,2,0,0,255] t1_subset
s: 333
i: [1,342,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 342
d: [2,889,1,2,1,1,40,315,558]
a: [3,895,1,2,1,889]
s: 895
i: [1,286,0,7,3,0,0,255] cc13_membered
s: 286
d: [2,890,1,2,1,1,40,315,518]
a: [3,900,1,2,1,890]
s: 900
i: [1,295,0,7,3,0,0,255] cc14_membered
s: 295
i: [1,298,0,7,3,0,0,255] cc17_membered
s: 298
d: [2,930,1,2,1,1,40,317,508]
a: [3,935,1,2,1,930]
s: 935
i: [1,335,0,7,2,0,0,255] t3_subset
s: 335
d: [2,948,1,4,2,1,40,324,635]
d: [2,950,1,2,1,2,21,948,640]
a: [3,954,1,2,1,950]
s: 954
i: [1,336,0,7,3,0,0,1] t8_boole
s: 336
i: [1,287,0,8,3,0,0,255] cc19_membered
s: 287
d: [2,949,1,4,2,1,40,324,660]
d: [2,956,1,2,1,2,21,949,665]
a: [3,960,1,2,1,956]
s: 960
i: [1,288,0,8,3,0,0,255] cc20_membered
s: 288
d: [2,962,1,4,2,1,40,326,635]
d: [2,964,1,2,1,2,21,962,640]
a: [3,968,1,2,1,964]
s: 968
i: [1,291,0,8,3,0,0,255] cc22_membered
s: 291
i: [1,292,0,8,3,0,0,255] cc23_membered
s: 292
d: [2,963,1,4,2,1,40,326,660]
d: [2,970,1,2,1,2,21,963,665]
a: [3,974,1,2,1,970]
s: 974
i: [1,296,0,8,3,0,0,255] cc21_membered
s: 296
d: [2,984,1,4,2,1,40,327,650]
d: [2,986,1,2,1,2,21,984,655]
a: [3,990,1,2,1,986]
s: 990
i: [1,299,0,8,3,0,0,3] cc24_membered
s: 299
i: [1,302,0,8,2,0,0,1] t12_taylor_1
s: 302
d: [2,985,1,4,2,1,40,327,670]
d: [2,992,1,2,1,2,21,985,675]
a: [3,996,1,2,1,992]
s: 996
i: [1,319,0,8,4,0,0,255] cc4_xxreal_0
s: 319
d: [2,1184,2,6,3,1,40,319,789]
a: [3,1209,2,2,1,1184]
s: 1209
d: [2,1245,3,0,0,3,52,718,383,1209,336]
e: 1245
d: [2,1246,3,0,0,3,52,718,383,1209,336]
e: 1246
d: [2,1247,3,0,0,3,52,383,1209,858]
e: 1247
r: 1209
i: [1,321,0,8,4,0,0,255] cc6_xxreal_0
s: 321
i: [1,322,0,8,4,0,0,255] cc8_xxreal_0
s: 322
a: [3,1208,2,2,1,1184]
s: 1208
i: [1,334,0,8,3,0,0,255] t2_subset
s: 334
a: [3,1213,2,2,1,1184]
s: 1213
i: [1,347,0,8,3,0,0,2] dt_k5_power
s: 347
d: [2,1181,0,8,2,1,22,302,453]
s: 1181
d: [2,1322,1,7,2,2,44,1181,378]
d: [2,1324,1,5,1,3,21,1322,388]
a: [3,1328,1,5,1,1324]
s: 1328
i: [1,339,0,9,3,0,0,255] t5_subset
s: 339
d: [2,1022,1,4,2,1,40,286,731]
a: [3,1034,1,2,1,1022]
s: 1034
i: [1,338,0,10,3,0,0,255] t4_subset
s: 338
i: [1,349,0,10,4,0,0,4] connectedness_r1_xxreal_0
s: 349
a: [3,1039,1,2,1,1022]
s: 1039
i: [1,306,0,11,5,0,0,255] t2_real
s: 306
d: [2,1357,1,8,4,1,40,306,488]
d: [2,1369,1,6,3,2,21,1357,900]
d: [2,1370,1,4,2,3,21,1369,856]
a: [3,1373,1,2,1,1370]
s: 1373
i: [1,307,0,11,5,0,0,3] t1_real
s: 307
i: [1,308,0,11,5,0,0,255] t8_real
s: 308
a: [3,1378,1,2,1,1370]
s: 1378
i: [1,309,0,11,5,0,0,255] t7_real
s: 309
d: [2,1029,1,4,2,1,40,286,518]
d: [2,1053,1,2,1,2,21,1029,749]
a: [3,1057,1,2,1,1053]
s: 1057
i: [1,310,0,11,5,0,0,255] t3_real
s: 310
i: [1,311,0,11,5,0,0,3] t4_real
s: 311
d: [2,1062,1,4,2,1,40,295,731]
a: [3,1075,1,2,1,1062]
s: 1075
i: [1,303,0,13,4,0,0,2] t51_power
s: 303
d: [2,1418,2,11,4,1,44,303,1328]
d: [2,1422,2,8,3,2,21,1418,383]
d: [2,1423,2,5,2,3,21,1422,712]
d: [2,1424,2,3,1,4,21,1423,413]
a: [3,1428,2,3,1,1424]
s: 1428
r: 1373
r: 1378
i: [1,304,0,13,6,0,0,255] t5_real
s: 304
i: [1,305,0,13,6,0,0,255] t6_real
s: 305
a: [3,1374,1,2,1,1370]
s: 1374
d: [2,1492,3,0,0,3,52,863,553,1374,326]
e: 1492
d: [2,1493,2,4,2,1,40,1374,326]
d: [2,1496,2,2,1,2,21,1493,553]
d: [2,1497,2,0,0,3,21,1496,863]
e: 1497
r: 1374
i: [1,346,0,14,4,0,0,2] existence_m2_subset_1
s: 346
s: 1373
i: [1,343,0,15,5,0,0,2] dt_m2_subset_1
s: 343
i: [1,344,0,15,5,0,0,2] redefinition_m2_subset_1
s: 344
s: 1378
i: [1,345,0,15,5,0,0,2] redefinition_m2_subset_1
s: 345
d: [2,1434,3,8,4,1,40,1428,307]
d: [2,1437,3,6,3,2,21,1434,1075]
d: [2,1438,3,4,2,3,21,1437,413]
a: [3,1442,3,2,1,1438]
s: 1442
d: [2,1524,4,0,0,3,52,1039,718,1442,327]
e: 1524
d: [2,1525,4,4,2,1,40,1442,327]
d: [2,1528,4,2,1,2,21,1525,718]
d: [2,1529,4,0,0,3,21,1528,1039]
e: 1529
d: [2,1141,2,2,1,1,40,935,294]
d: [2,1407,1,8,4,1,40,308,707]
d: [2,1412,1,6,3,2,21,1407,900]
d: [2,1413,1,4,2,3,21,1412,413]
d: [2,1420,2,11,4,1,44,1328,303]
d: [2,1430,2,8,3,2,21,1420,383]
d: [2,1431,2,5,2,3,21,1430,712]
d: [2,1432,2,3,1,4,21,1431,413]
f: 1529,1497,1438,1432,1413,1370,1324,1141,1062,1022,930,890,859,852,727,714,703,702,277,276,271,269,267,262,261,260,256,255,253,252,245,246,247
% Refutation found. Thanks to Tanya!
% SZS status Theorem for sin_cos7__t29_sin_cos7
% SZS output start Proof for sin_cos7__t29_sin_cos7
1. ! [X0] : (v1_xreal_0(X0) => ~(k25_sin_cos(X0) = 1 & k6_numbers != X0)) [input t29_sin_cos7]
2. ~! [X0] : (v1_xreal_0(X0) => ~(k25_sin_cos(X0) = 1 & k6_numbers != X0)) [negated conjecture 1]
6. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xxreal_0(X1))) [input cc13_membered]
7. ! [X0] : (v3_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xreal_0(X1))) [input cc14_membered]
14. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xreal_0(X0)) [input cc1_xreal_0]
27. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
28. ! [X0] : (v3_membered(X0) => v2_membered(X0)) [input cc4_membered]
29. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
32. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
35. ! [X0] : (m1_subset_1(X0,k1_zfmisc_1(k1_numbers)) => v3_membered(X0)) [input cc8_membered]
43. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [input dt_k5_numbers]
46. v1_xreal_0(k7_power) [input dt_k7_power]
52. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
56. ! [X0] : (v1_xreal_0(X0) => v1_xreal_0(k25_sin_cos(X0))) [input fc5_sin_cos]
58. 1 != k8_power [input l1_sin_cos7]
72. k5_numbers = k4_ordinal1 [input redefinition_k5_numbers]
73. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
74. k7_power = k8_power [input redefinition_k8_power]
83. r1_xxreal_0(1,2) [input rqLessOrEqual__r1_xxreal_0__r1_r2]
90. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
92. m1_subset_1(2,k1_numbers) & m1_subset_1(2,k5_numbers) & m2_subset_1(2,k1_numbers,k5_numbers) & v2_xxreal_0(2) [input spc2_numerals]
93. ~r1_xxreal_0(k8_power,2) [input t11_taylor_1]
94. ! [X0] : (v1_xreal_0(X0) => k5_power(k8_power,k25_sin_cos(X0)) = X0) [input t12_taylor_1]
95. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
96. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
98. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
104. ! [X0] : (v1_xreal_0(X0) => ~(k6_numbers != k5_power(X0,1) & 1 != X0 & ~r1_xxreal_0(X0,k6_numbers))) [input t51_power]
112. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
133. ? [X0] : ((k25_sin_cos(X0) = 1 & k6_numbers != X0) & v1_xreal_0(X0)) [ennf transformation 2]
134. ? [X0] : (k25_sin_cos(X0) = 1 & k6_numbers != X0 & v1_xreal_0(X0)) [flattening 133]
137. ! [X0] : (! [X1] : (v1_xxreal_0(X1) | ~m1_subset_1(X1,X0)) | ~v2_membered(X0)) [ennf transformation 6]
145. ! [X0] : (v2_membered(X0) | ~v3_membered(X0)) [ennf transformation 28]
146. ! [X0] : (! [X1] : (v1_xreal_0(X1) | ~m1_subset_1(X1,X0)) | ~v3_membered(X0)) [ennf transformation 7]
151. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 29]
152. ! [X0] : (v1_xreal_0(k25_sin_cos(X0)) | ~v1_xreal_0(X0)) [ennf transformation 56]
153. ! [X0] : (k5_power(k8_power,k25_sin_cos(X0)) = X0 | ~v1_xreal_0(X0)) [ennf transformation 94]
154. ! [X0] : ((k6_numbers = k5_power(X0,1) | 1 = X0 | r1_xxreal_0(X0,k6_numbers)) | ~v1_xreal_0(X0)) [ennf transformation 104]
155. ! [X0] : (k6_numbers = k5_power(X0,1) | 1 = X0 | r1_xxreal_0(X0,k6_numbers) | ~v1_xreal_0(X0)) [flattening 154]
160. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 98]
161. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 160]
162. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 96]
163. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 162]
164. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 112]
165. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 164]
175. ! [X0] : (v1_xreal_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 14]
177. ! [X0] : (v3_membered(X0) | ~m1_subset_1(X0,k1_zfmisc_1(k1_numbers))) [ennf transformation 35]
184. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 32]
185. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 184]
186. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 27]
187. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 186]
212. ? [X0] : (k25_sin_cos(X0) = 1 & k6_numbers != X0 & v1_xreal_0(X0)) => (1 = k25_sin_cos(sK0) & k6_numbers != sK0 & v1_xreal_0(sK0)) [choice axiom]
213. 1 = k25_sin_cos(sK0) & k6_numbers != sK0 & v1_xreal_0(sK0) [skolemisation 134,212]
245. v1_xreal_0(sK0) [cnf transformation 213] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
246. k6_numbers != sK0 [cnf transformation 213] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
247. 1 = k25_sin_cos(sK0) [cnf transformation 213] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
252. v1_xreal_0(k7_power) [cnf transformation 46] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
253. v1_xboole_0(k1_xboole_0) [cnf transformation 52] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
255. 1 != k8_power [cnf transformation 58] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
256. ~r1_xxreal_0(k8_power,2) [cnf transformation 93] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
260. k7_power = k8_power [cnf transformation 74] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
261. k6_numbers = k1_xboole_0 [cnf transformation 73] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
262. k5_numbers = k4_ordinal1 [cnf transformation 72] {a:0,w:3,wCS:1200,sine:3,thAx:0,allAx:1,thDist:-1}
267. r1_xxreal_0(1,2) [cnf transformation 83] {a:0,w:3,wCS:1200,sine:3,thAx:0,allAx:1,thDist:-1}
269. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 95] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
271. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [cnf transformation 43] {a:0,w:4,wCS:1600,sine:2,thAx:0,allAx:1,thDist:-1}
276. m1_subset_1(2,k1_numbers) [cnf transformation 92] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
277. v2_xxreal_0(1) [cnf transformation 90] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
286. ~m1_subset_1(X1,X0) | v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 137] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
294. ~v3_membered(X0) | v2_membered(X0) [cnf transformation 145] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
295. ~m1_subset_1(X1,X0) | v1_xreal_0(X1) | ~v3_membered(X0) [cnf transformation 146] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
300. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 151] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
301. v1_xreal_0(k25_sin_cos(X0)) | ~v1_xreal_0(X0) [cnf transformation 152] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
302. k5_power(k8_power,k25_sin_cos(X0)) = X0 | ~v1_xreal_0(X0) [cnf transformation 153] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
303. k6_numbers = k5_power(X0,1) | 1 = X0 | r1_xxreal_0(X0,k6_numbers) | ~v1_xreal_0(X0) [cnf transformation 155] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
306. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 161] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
307. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 163] {a:0,w:11,wCS:4400,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
308. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 165] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
315. ~m1_subset_1(X0,k1_numbers) | v1_xreal_0(X0) [cnf transformation 175] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
317. ~m1_subset_1(X0,k1_zfmisc_1(k1_numbers)) | v3_membered(X0) [cnf transformation 177] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
326. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 185] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
327. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 187] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
376. 1 <=> 1 = k25_sin_cos(sK0) [avatar definition]
378. 1 = k25_sin_cos(sK0) <- (1) [avatar component clause 376] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
379. 1 [avatar split clause 247,376]
381. 2 <=> k6_numbers = sK0 [avatar definition]
383. k6_numbers != sK0 <- (~2) [avatar component clause 381] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
384. ~2 [avatar split clause 246,381]
386. 3 <=> v1_xreal_0(sK0) [avatar definition]
388. v1_xreal_0(sK0) <- (3) [avatar component clause 386] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
389. 3 [avatar split clause 245,386]
411. 8 <=> v1_xreal_0(k7_power) [avatar definition]
413. v1_xreal_0(k7_power) <- (8) [avatar component clause 411] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
414. 8 [avatar split clause 252,411]
416. 9 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
418. v1_xboole_0(k1_xboole_0) <- (9) [avatar component clause 416] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
419. 9 [avatar split clause 253,416]
426. 11 <=> 1 = k8_power [avatar definition]
428. 1 != k8_power <- (~11) [avatar component clause 426] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
429. ~11 [avatar split clause 255,426]
431. 12 <=> r1_xxreal_0(k8_power,2) [avatar definition]
433. ~r1_xxreal_0(k8_power,2) <- (~12) [avatar component clause 431] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
434. ~12 [avatar split clause 256,431]
451. 16 <=> k7_power = k8_power [avatar definition]
453. k7_power = k8_power <- (16) [avatar component clause 451] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
454. 16 [avatar split clause 260,451]
456. 17 <=> k6_numbers = k1_xboole_0 [avatar definition]
458. k6_numbers = k1_xboole_0 <- (17) [avatar component clause 456] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
459. 17 [avatar split clause 261,456]
461. 18 <=> k5_numbers = k4_ordinal1 [avatar definition]
463. k5_numbers = k4_ordinal1 <- (18) [avatar component clause 461] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
464. 18 [avatar split clause 262,461]
486. 23 <=> r1_xxreal_0(1,2) [avatar definition]
488. r1_xxreal_0(1,2) <- (23) [avatar component clause 486] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
489. 23 [avatar split clause 267,486]
496. 25 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
498. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (25) [avatar component clause 496] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
499. 25 [avatar split clause 269,496]
506. 27 <=> m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [avatar definition]
508. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) <- (27) [avatar component clause 506] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
509. 27 [avatar split clause 271,506]
516. 29 <=> m1_subset_1(2,k1_numbers) [avatar definition]
518. m1_subset_1(2,k1_numbers) <- (29) [avatar component clause 516] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
519. 29 [avatar split clause 276,516]
551. 36 <=> v2_xxreal_0(1) [avatar definition]
553. v2_xxreal_0(1) <- (36) [avatar component clause 551] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
554. 36 [avatar split clause 277,551]
702. ~r1_xxreal_0(k7_power,2) <- (~12, 16) [backward demodulation 433,453] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
703. 1 != k7_power <- (~11, 16) [superposition 428,453] {a:1,w:3,wCS:1200,sine:255,thAx:0,allAx:2,thDist:-2}
705. 67 <=> r1_xxreal_0(k7_power,2) [avatar definition]
707. ~r1_xxreal_0(k7_power,2) <- (~67) [avatar component clause 705] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
708. ~67 | 12 | ~16 [avatar split clause 702,451,431,705]
710. 68 <=> 1 = k7_power [avatar definition]
712. 1 != k7_power <- (~68) [avatar component clause 710] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
713. ~68 | 11 | ~16 [avatar split clause 703,451,426,710]
714. v1_xboole_0(k6_numbers) <- (9, 17) [backward demodulation 418,458] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
716. 69 <=> v1_xboole_0(k6_numbers) [avatar definition]
718. v1_xboole_0(k6_numbers) <- (69) [avatar component clause 716] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
719. 69 | ~9 | ~17 [avatar split clause 714,456,416,716]
726. m1_subset_1(k1_xboole_0,k5_numbers) <- (18, 25) [forward demodulation 498,463] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
727. m1_subset_1(k6_numbers,k5_numbers) <- (17, 18, 25) [forward demodulation 726,458] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
729. 71 <=> m1_subset_1(k6_numbers,k5_numbers) [avatar definition]
731. m1_subset_1(k6_numbers,k5_numbers) <- (71) [avatar component clause 729] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
732. 71 | ~17 | ~18 | ~25 [avatar split clause 727,496,461,456,729]
851. v1_xreal_0(1) | ~v1_xreal_0(sK0) <- (1) [superposition 301,378] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
852. v1_xreal_0(1) <- (1, 3) [subsumption resolution 851,388] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
854. 91 <=> v1_xreal_0(1) [avatar definition]
856. v1_xreal_0(1) <- (91) [avatar component clause 854] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
857. 91 | ~1 | ~3 [avatar split clause 852,386,376,854]
859. v1_xxreal_0(1) <- (91) [resolution 856,300] {a:2,w:2,wCS:200,goal:1,sine:4,thAx:0,allAx:4,thDist:-4}
861. 92 <=> v1_xxreal_0(1) [avatar definition]
863. v1_xxreal_0(1) <- (92) [avatar component clause 861] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
864. 92 | ~91 [avatar split clause 859,854,861]
890. v1_xreal_0(2) <- (29) [resolution 315,518] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
898. 97 <=> v1_xreal_0(2) [avatar definition]
900. v1_xreal_0(2) <- (97) [avatar component clause 898] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
901. 97 | ~29 [avatar split clause 890,516,898]
930. v3_membered(k5_numbers) <- (27) [resolution 317,508] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
933. 102 <=> v3_membered(k5_numbers) [avatar definition]
935. v3_membered(k5_numbers) <- (102) [avatar component clause 933] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
936. 102 | ~27 [avatar split clause 930,506,933]
1022. v1_xxreal_0(k6_numbers) | ~v2_membered(k5_numbers) <- (71) [resolution 286,731] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1033. 116 <=> v2_membered(k5_numbers) [avatar definition]
1037. 117 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1039. v1_xxreal_0(k6_numbers) <- (117) [avatar component clause 1037] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1040. ~116 | 117 | ~71 [avatar split clause 1022,729,1037,1033]
1062. v1_xreal_0(k6_numbers) | ~v3_membered(k5_numbers) <- (71) [resolution 295,731] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1073. 121 <=> v1_xreal_0(k6_numbers) [avatar definition]
1075. v1_xreal_0(k6_numbers) <- (121) [avatar component clause 1073] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1076. ~102 | 121 | ~71 [avatar split clause 1062,729,1073,933]
1141. v2_membered(k5_numbers) <- (102) [resolution 935,294] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1143. 116 | ~102 [avatar split clause 1141,933,1033]
1181. k5_power(k7_power,k25_sin_cos(X0)) = X0 | ~v1_xreal_0(X0) <- (16) [forward demodulation 302,453] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1198. 135 <=> v2_xxreal_0(k7_power) [avatar definition]
1322. sK0 = k5_power(k7_power,1) | ~v1_xreal_0(sK0) <- (1, 16) [superposition 1181,378] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1324. sK0 = k5_power(k7_power,1) <- (1, 3, 16) [subsumption resolution 1322,388] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1326. 158 <=> sK0 = k5_power(k7_power,1) [avatar definition]
1328. sK0 = k5_power(k7_power,1) <- (158) [avatar component clause 1326] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1329. 158 | ~1 | ~3 | ~16 [avatar split clause 1324,451,386,376,1326]
1357. ~v3_xxreal_0(2) | v3_xxreal_0(1) | ~v1_xreal_0(2) | ~v1_xreal_0(1) <- (23) [resolution 306,488] {a:1,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
1369. ~v3_xxreal_0(2) | v3_xxreal_0(1) | ~v1_xreal_0(1) <- (23, 97) [subsumption resolution 1357,900] {a:1,w:6,wCS:2400,sine:255,thAx:0,allAx:2,thDist:-2}
1370. ~v3_xxreal_0(2) | v3_xxreal_0(1) <- (23, 91, 97) [subsumption resolution 1369,856] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1372. 161 <=> v3_xxreal_0(1) [avatar definition]
1374. v3_xxreal_0(1) <- (161) [avatar component clause 1372] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1376. 162 <=> v3_xxreal_0(2) [avatar definition]
1379. 161 | ~162 | ~23 | ~91 | ~97 [avatar split clause 1370,898,854,486,1376,1372]
1407. v3_xxreal_0(2) | v2_xxreal_0(k7_power) | ~v1_xreal_0(2) | ~v1_xreal_0(k7_power) <- (~67) [resolution 308,707] {a:1,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
1412. v3_xxreal_0(2) | v2_xxreal_0(k7_power) | ~v1_xreal_0(k7_power) <- (~67, 97) [subsumption resolution 1407,900] {a:1,w:6,wCS:2400,sine:255,thAx:0,allAx:2,thDist:-2}
1413. v3_xxreal_0(2) | v2_xxreal_0(k7_power) <- (8, ~67, 97) [subsumption resolution 1412,413] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1414. 135 | 162 | ~8 | 67 | ~97 [avatar split clause 1413,898,705,411,1376,1198]
1420. k6_numbers = sK0 | 1 = k7_power | r1_xxreal_0(k7_power,k6_numbers) | ~v1_xreal_0(k7_power) <- (158) [superposition 1328,303] {a:2,w:11,wCS:1100,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1426. 164 <=> r1_xxreal_0(k7_power,k6_numbers) [avatar definition]
1428. r1_xxreal_0(k7_power,k6_numbers) <- (164) [avatar component clause 1426] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1430. 1 = k7_power | r1_xxreal_0(k7_power,k6_numbers) | ~v1_xreal_0(k7_power) <- (~2, 158) [subsumption resolution 1420,383] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1431. r1_xxreal_0(k7_power,k6_numbers) | ~v1_xreal_0(k7_power) <- (~2, ~68, 158) [subsumption resolution 1430,712] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1432. r1_xxreal_0(k7_power,k6_numbers) <- (~2, 8, ~68, 158) [subsumption resolution 1431,413] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1433. 164 | 2 | ~8 | 68 | ~158 [avatar split clause 1432,1326,710,411,381,1426]
1434. ~v2_xxreal_0(k7_power) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k7_power) <- (164) [resolution 1428,307] {a:3,w:8,wCS:800,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
1437. ~v2_xxreal_0(k7_power) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k7_power) <- (121, 164) [subsumption resolution 1434,1075] {a:3,w:6,wCS:600,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
1438. ~v2_xxreal_0(k7_power) | v2_xxreal_0(k6_numbers) <- (8, 121, 164) [subsumption resolution 1437,413] {a:3,w:4,wCS:400,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
1440. 165 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1442. v2_xxreal_0(k6_numbers) <- (165) [avatar component clause 1440] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1443. 165 | ~135 | ~8 | ~121 | ~164 [avatar split clause 1438,1426,1073,411,1198,1440]
1493. ~v2_xxreal_0(1) | ~v1_xxreal_0(1) <- (161) [resolution 1374,326] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1496. ~v1_xxreal_0(1) <- (36, 161) [subsumption resolution 1493,553] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1497. $false <- (36, 92, 161) [subsumption resolution 1496,863] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1498. ~36 | ~92 | ~161 [avatar contradiction clause 1497]
1525. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (165) [resolution 1442,327] {a:4,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1528. ~v1_xxreal_0(k6_numbers) <- (69, 165) [subsumption resolution 1525,718] {a:4,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1529. $false <- (69, 117, 165) [subsumption resolution 1528,1039] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1530. ~69 | ~117 | ~165 [avatar contradiction clause 1529]
1531. $false [avatar sat refutation 379,384,389,414,419,429,434,454,459,464,489,499,509,519,554,708,713,719,732,857,864,901,936,1040,1076,1143,1329,1379,1414,1433,1443,1498,1530] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:762,thDist:-762}
% SZS output end Proof for sin_cos7__t29_sin_cos7
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 111
% >>> Preprocessing
% Introduced skolems: 16
% Pure predicates: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 130
% Generated clauses: 595
% Active clauses: 162
% Passive clauses: 260
% Final active clauses: 155
% Final passive clauses: 81
% Main loop iterations started: 174
% 
% >>> Simplifying Inferences
% Duplicate literals: 32
% Fw subsumption resolutions: 46
% Fw demodulations: 5
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Forward subsumptions: 77
% 
% >>> Generating Inferences
% Binary resolution: 206
% Unit resulting resolution: 5
% Factoring: 1
% Forward superposition: 7
% Backward superposition: 9
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 39
% Split components: 85
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 177
% SAT solver unit clauses: 68
% SAT solver binary clauses: 40
% 
% Memory used [KB]: 5756
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% SAT solver time: 0.001 s
% minimizing solver time: 0.001 s
% congruence closure: 0.002 s ( own 0.001 s ) 
% splitting component index maintenance: 0.001 s
% forward subsumption: 0.001 s
% other: 0.007 s ( own 0.001 s ) 

