i: [1,31,0,11,4,1,0,0] t19_gate_1
a: [3,56,0,2,1,31]
s: 56
a: [3,43,0,5,1,31]
s: 43
i: [1,32,0,2,1,0,0,255] fc1_xboole_0
a: [3,64,0,2,1,32]
s: 64
i: [1,39,0,2,1,0,0,1] rc2_xboole_0
a: [3,69,0,2,1,39]
s: 69
i: [1,40,0,2,1,0,0,1] rc1_xboole_0
a: [3,74,0,2,1,40]
s: 74
i: [1,33,0,5,2,0,0,255] fc2_gate_1
s: 33
i: [1,34,0,5,2,0,0,255] fc1_gate_1
s: 34
i: [1,36,0,8,2,0,0,255] d13_gate_1
s: 36
a: [3,47,0,2,1,31]
s: 47
i: [1,37,0,8,2,0,0,255] d13_gate_1
s: 37
a: [3,51,0,2,1,31]
s: 51
i: [1,38,0,8,2,0,0,255] d13_gate_1
s: 38
d: [2,83,1,4,2,1,44,43,38]
d: [2,84,1,2,1,2,21,83,64]
d: [2,85,1,0,0,3,21,84,56]
e: 85
r: 43
r: 56
a: [3,55,0,2,1,31]
s: 55
a: [3,44,0,5,1,31]
s: 44
i: [1,35,0,13,4,0,0,255] d13_gate_1
s: 35
d: [2,96,1,9,4,1,44,44,35]
d: [2,97,1,7,3,2,21,96,55]
d: [2,98,1,5,2,3,21,97,51]
d: [2,99,1,3,1,4,21,98,47]
a: [3,103,1,3,1,99]
s: 103
d: [2,105,2,0,0,2,52,64,103,34]
e: 105
d: [2,106,2,2,1,1,40,103,34]
d: [2,108,2,0,0,2,21,106,64]
e: 108
i: [1,30,0,7,2,1,0,0] t19_gate_1
i: [1,29,0,7,2,1,0,0] t19_gate_1
i: [1,28,0,7,2,1,0,0] t19_gate_1
d: [2,77,1,4,2,1,44,43,36]
d: [2,78,1,2,1,2,21,77,64]
d: [2,80,1,4,2,1,44,43,37]
d: [2,81,1,2,1,2,21,80,64]
f: 108,99,85,81,78,32,28,29,30,31
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t19_gate_1
% SZS output start Proof for gate_1__t19_gate_1
1. ! [X0] : ! [X1] : ! [X2] : (~v1_xboole_0(k13_gate_1(X0,X1,X2)) <=> (v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [input t19_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : (~v1_xboole_0(k13_gate_1(X0,X1,X2)) <=> (v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ((~(v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k13_gate_1(X0,X1,X2) = k1_xboole_0) & ~(k13_gate_1(X0,X1,X2) != k1_gate_1(k1_xboole_0) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [input d13_gate_1]
14. ~! [X0,X1,X2] : (~v1_xboole_0(k13_gate_1(X0,X1,X2)) <=> (v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [flattening 2]
15. ! [X0,X1,X2] : ((~(v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) => k13_gate_1(X0,X1,X2) = k1_xboole_0) & ~(k13_gate_1(X0,X1,X2) != k1_gate_1(k1_xboole_0) & v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [flattening 13]
16. ? [X0,X1,X2] : (~v1_xboole_0(k13_gate_1(X0,X1,X2)) <~> (v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2] : ((k13_gate_1(X0,X1,X2) = k1_xboole_0 | (v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0))) & (k13_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 15]
20. ? [X0,X1,X2] : (((~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) | v1_xboole_0(k13_gate_1(X0,X1,X2))) & ((v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) | ~v1_xboole_0(k13_gate_1(X0,X1,X2)))) [nnf transformation 16]
21. ? [X0,X1,X2] : ((~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) | v1_xboole_0(k13_gate_1(X0,X1,X2))) & ((v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) | ~v1_xboole_0(k13_gate_1(X0,X1,X2)))) [flattening 20]
22. ? [X0,X1,X2] : ((~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) | v1_xboole_0(k13_gate_1(X0,X1,X2))) & ((v1_xboole_0(X2) & v1_xboole_0(X1) & v1_xboole_0(X0)) | ~v1_xboole_0(k13_gate_1(X0,X1,X2)))) => ((~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | v1_xboole_0(k13_gate_1(sK0,sK1,sK2))) & ((v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0)) | ~v1_xboole_0(k13_gate_1(sK0,sK1,sK2)))) [choice axiom]
23. (~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | v1_xboole_0(k13_gate_1(sK0,sK1,sK2))) & ((v1_xboole_0(sK2) & v1_xboole_0(sK1) & v1_xboole_0(sK0)) | ~v1_xboole_0(k13_gate_1(sK0,sK1,sK2))) [skolemisation 21,22]
28. v1_xboole_0(sK0) | ~v1_xboole_0(k13_gate_1(sK0,sK1,sK2)) [cnf transformation 23] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. v1_xboole_0(sK1) | ~v1_xboole_0(k13_gate_1(sK0,sK1,sK2)) [cnf transformation 23] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. v1_xboole_0(sK2) | ~v1_xboole_0(k13_gate_1(sK0,sK1,sK2)) [cnf transformation 23] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | v1_xboole_0(k13_gate_1(sK0,sK1,sK2)) [cnf transformation 23] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
34. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
35. k13_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X2) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
36. k13_gate_1(X0,X1,X2) = k1_xboole_0 | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
37. k13_gate_1(X0,X1,X2) = k1_xboole_0 | v1_xboole_0(X1) [cnf transformation 19] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. k13_gate_1(X0,X1,X2) = k1_xboole_0 | v1_xboole_0(X2) [cnf transformation 19] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. 1 <=> v1_xboole_0(k13_gate_1(sK0,sK1,sK2)) [avatar definition]
43. ~v1_xboole_0(k13_gate_1(sK0,sK1,sK2)) <- (~1) [avatar component clause 42] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
44. v1_xboole_0(k13_gate_1(sK0,sK1,sK2)) <- (1) [avatar component clause 42] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
46. 2 <=> v1_xboole_0(sK0) [avatar definition]
47. v1_xboole_0(sK0) <- (2) [avatar component clause 46] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
50. 3 <=> v1_xboole_0(sK1) [avatar definition]
51. v1_xboole_0(sK1) <- (3) [avatar component clause 50] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
54. 4 <=> v1_xboole_0(sK2) [avatar definition]
55. v1_xboole_0(sK2) <- (4) [avatar component clause 54] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. ~v1_xboole_0(sK2) <- (~4) [avatar component clause 54] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
57. 1 | ~2 | ~3 | ~4 [avatar split clause 31,54,50,46,42]
58. ~1 | 4 [avatar split clause 30,54,42]
59. ~1 | 3 [avatar split clause 29,50,42]
60. ~1 | 2 [avatar split clause 28,46,42]
62. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
64. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 62] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
65. 5 [avatar split clause 32,62]
77. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK0) <- (~1) [superposition 43,36] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
78. v1_xboole_0(sK0) <- (~1, 5) [subsumption resolution 77,64] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
79. 2 | 1 | ~5 [avatar split clause 78,62,42,46]
80. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK1) <- (~1) [superposition 43,37] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
81. v1_xboole_0(sK1) <- (~1, 5) [subsumption resolution 80,64] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. 3 | 1 | ~5 [avatar split clause 81,62,42,50]
83. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK2) <- (~1) [superposition 43,38] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
84. v1_xboole_0(sK2) <- (~1, 5) [subsumption resolution 83,64] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
85. $false <- (~1, ~4, 5) [subsumption resolution 84,56] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. 1 | 4 | ~5 [avatar contradiction clause 85]
96. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1) [superposition 44,35] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
97. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1, 4) [subsumption resolution 96,55] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK0) <- (1, 3, 4) [subsumption resolution 97,51] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
99. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (1, 2, 3, 4) [subsumption resolution 98,47] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. 8 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
103. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (8) [avatar component clause 101] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. 8 | ~1 | ~2 | ~3 | ~4 [avatar split clause 99,54,50,46,42,101]
106. ~v1_xboole_0(k1_xboole_0) <- (8) [resolution 103,34] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
108. $false <- (5, 8) [subsumption resolution 106,64] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
109. ~5 | ~8 [avatar contradiction clause 108]
110. $false [avatar sat refutation 57,58,59,60,65,79,82,86,104,109] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:58,thDist:-58}
% SZS output end Proof for gate_1__t19_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 47
% Active clauses: 16
% Passive clauses: 16
% Final active clauses: 14
% Main loop iterations started: 16
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 8
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 6
% Backward superposition: 7
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 10
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 3
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

