i: [1,38,0,2,1,1,0,0] t1_hessenbe
a: [3,96,0,2,1,38]
s: 96
i: [1,37,0,2,1,1,0,0] t1_hessenbe
a: [3,101,0,2,1,37]
s: 101
i: [1,36,0,2,1,1,0,0] t1_hessenbe
a: [3,106,0,2,1,36]
s: 106
i: [1,35,0,2,1,1,0,0] t1_hessenbe
a: [3,111,0,2,1,35]
s: 111
i: [1,41,0,4,1,1,0,0] t1_hessenbe
a: [3,81,0,4,1,41]
s: 81
i: [1,40,0,4,1,1,0,0] t1_hessenbe
a: [3,86,0,4,1,40]
s: 86
i: [1,39,0,4,1,1,0,0] t1_hessenbe
a: [3,91,0,4,1,39]
s: 91
i: [1,43,0,25,5,1,0,0] t1_hessenbe
a: [3,55,0,5,1,43]
s: 55
i: [1,42,0,5,1,1,0,0] t1_hessenbe
a: [3,76,0,5,1,42]
s: 76
i: [1,47,0,2,1,0,0,1] existence_l1_collsp
a: [3,116,0,2,1,47]
s: 116
i: [1,51,0,2,1,0,0,1] rc2_anproj_2
a: [3,121,0,2,1,51]
s: 121
i: [1,50,0,2,1,0,0,1] rc2_anproj_2
a: [3,126,0,2,1,50]
s: 126
i: [1,49,0,2,1,0,0,1] rc2_anproj_2
a: [3,131,0,2,1,49]
s: 131
i: [1,48,0,2,1,0,0,1] rc2_anproj_2
a: [3,136,0,2,1,48]
s: 136
i: [1,46,0,4,1,0,0,1] existence_m1_subset_1
s: 46
i: [1,44,0,30,9,0,0,1] t4_collsp
s: 44
a: [3,62,0,5,1,43]
s: 62
i: [1,45,0,30,9,0,0,1] t4_collsp
s: 45
d: [2,148,1,0,0,9,52,111,106,101,96,86,91,62,55,81,45]
e: 148
d: [2,150,1,25,8,1,40,45,62]
d: [2,159,1,20,7,2,21,150,55]
d: [2,160,1,16,6,3,21,159,81]
d: [2,161,1,12,5,4,21,160,91]
d: [2,162,1,8,4,5,21,161,86]
d: [2,163,1,6,3,6,21,162,96]
d: [2,164,1,4,2,7,21,163,101]
d: [2,165,1,2,1,8,21,164,106]
d: [2,166,1,0,0,9,21,165,111]
e: 166
r: 55
a: [3,54,0,5,1,43]
s: 54
a: [3,59,0,5,1,43]
s: 59
a: [3,66,0,5,1,43]
s: 66
d: [2,187,1,0,0,9,52,111,106,101,96,91,81,86,59,66,44]
e: 187
d: [2,189,1,25,8,1,40,66,44]
d: [2,191,1,20,7,2,21,189,59]
d: [2,192,1,16,6,3,21,191,86]
d: [2,193,1,12,5,4,21,192,81]
d: [2,194,1,8,4,5,21,193,91]
d: [2,195,1,6,3,6,21,194,96]
d: [2,196,1,4,2,7,21,195,101]
d: [2,197,1,2,1,8,21,196,106]
d: [2,198,1,0,0,9,21,197,111]
e: 198
d: [2,138,1,25,8,1,40,44,76]
d: [2,139,1,21,7,2,21,138,81]
d: [2,140,1,17,6,3,21,139,86]
d: [2,141,1,13,5,4,21,140,91]
d: [2,142,1,11,4,5,21,141,96]
d: [2,143,1,9,3,6,21,142,101]
d: [2,144,1,7,2,7,21,143,106]
d: [2,145,1,5,1,8,21,144,111]
d: [2,168,1,21,7,2,21,150,81]
d: [2,169,1,17,6,3,21,168,91]
d: [2,170,1,13,5,4,21,169,86]
d: [2,171,1,11,4,5,21,170,96]
d: [2,172,1,9,3,6,21,171,101]
d: [2,173,1,7,2,7,21,172,106]
d: [2,174,1,5,1,8,21,173,111]
d: [2,149,1,25,8,1,40,45,76]
d: [2,152,1,21,7,2,21,149,81]
d: [2,153,1,17,6,3,21,152,86]
d: [2,154,1,13,5,4,21,153,91]
d: [2,155,1,11,4,5,21,154,96]
d: [2,156,1,9,3,6,21,155,101]
d: [2,157,1,7,2,7,21,156,106]
d: [2,158,1,5,1,8,21,157,111]
d: [2,178,1,25,8,1,40,54,44]
d: [2,179,1,21,7,2,21,178,91]
d: [2,180,1,17,6,3,21,179,81]
d: [2,181,1,13,5,4,21,180,86]
d: [2,182,1,11,4,5,21,181,96]
d: [2,183,1,9,3,6,21,182,101]
d: [2,184,1,7,2,7,21,183,106]
d: [2,185,1,5,1,8,21,184,111]
f: 198,185,158,174,145,35,36,37,38,39,40,41,42,43
% Refutation found. Thanks to Tanya!
% SZS status Theorem for hessenbe__t1_hessenbe
% SZS output start Proof for hessenbe__t1_hessenbe
1. ! [X0] : ((l1_collsp(X0) & v3_anproj_2(X0) & v2_anproj_2(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [input t1_hessenbe]
2. ~! [X0] : ((l1_collsp(X0) & v3_anproj_2(X0) & v2_anproj_2(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [negated conjecture 1]
11. ! [X0] : ((l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3))))))) [input t4_collsp]
14. ~! [X0] : ((l1_collsp(X0) & v2_anproj_2(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [pure predicate removal 2]
16. ~! [X0] : ((l1_collsp(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [pure predicate removal 14]
19. ~! [X0] : ((l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_collsp(X0,X1,X2,X3) => (r1_collsp(X0,X3,X2,X1) & r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3) & r1_collsp(X0,X3,X1,X2) & r1_collsp(X0,X2,X3,X1))))))) [pure predicate removal 16]
20. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~r1_collsp(X0,X3,X2,X1) | ~r1_collsp(X0,X1,X3,X2) | ~r1_collsp(X0,X2,X1,X3) | ~r1_collsp(X0,X3,X1,X2) | ~r1_collsp(X0,X2,X3,X1)) & r1_collsp(X0,X1,X2,X3)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0))) [ennf transformation 19]
21. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_collsp(X0,X3,X2,X1) | ~r1_collsp(X0,X1,X3,X2) | ~r1_collsp(X0,X2,X1,X3) | ~r1_collsp(X0,X3,X1,X2) | ~r1_collsp(X0,X2,X3,X1)) & r1_collsp(X0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) [flattening 20]
22. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3)) | ~r1_collsp(X0,X1,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0))) [ennf transformation 11]
23. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_collsp(X0,X1,X3,X2) & r1_collsp(X0,X2,X1,X3)) | ~r1_collsp(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0)) [flattening 22]
24. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_collsp(X0,X3,X2,X1) | ~r1_collsp(X0,X1,X3,X2) | ~r1_collsp(X0,X2,X1,X3) | ~r1_collsp(X0,X3,X1,X2) | ~r1_collsp(X0,X2,X3,X1)) & r1_collsp(X0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : ((~r1_collsp(sK0,X3,X2,X1) | ~r1_collsp(sK0,X1,X3,X2) | ~r1_collsp(sK0,X2,X1,X3) | ~r1_collsp(sK0,X3,X1,X2) | ~r1_collsp(sK0,X2,X3,X1)) & r1_collsp(sK0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_collsp(sK0) & v3_collsp(sK0) & v2_collsp(sK0) & ~v2_struct_0(sK0)) [choice axiom]
25. ? [X1] : (? [X2] : (? [X3] : ((~r1_collsp(sK0,X3,X2,X1) | ~r1_collsp(sK0,X1,X3,X2) | ~r1_collsp(sK0,X2,X1,X3) | ~r1_collsp(sK0,X3,X1,X2) | ~r1_collsp(sK0,X2,X3,X1)) & r1_collsp(sK0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : ((~r1_collsp(sK0,X3,X2,sK1) | ~r1_collsp(sK0,sK1,X3,X2) | ~r1_collsp(sK0,X2,sK1,X3) | ~r1_collsp(sK0,X3,sK1,X2) | ~r1_collsp(sK0,X2,X3,sK1)) & r1_collsp(sK0,sK1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
26. ? [X2] : (? [X3] : ((~r1_collsp(sK0,X3,X2,sK1) | ~r1_collsp(sK0,sK1,X3,X2) | ~r1_collsp(sK0,X2,sK1,X3) | ~r1_collsp(sK0,X3,sK1,X2) | ~r1_collsp(sK0,X2,X3,sK1)) & r1_collsp(sK0,sK1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : ((~r1_collsp(sK0,X3,sK2,sK1) | ~r1_collsp(sK0,sK1,X3,sK2) | ~r1_collsp(sK0,sK2,sK1,X3) | ~r1_collsp(sK0,X3,sK1,sK2) | ~r1_collsp(sK0,sK2,X3,sK1)) & r1_collsp(sK0,sK1,sK2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
27. ? [X3] : ((~r1_collsp(sK0,X3,sK2,sK1) | ~r1_collsp(sK0,sK1,X3,sK2) | ~r1_collsp(sK0,sK2,sK1,X3) | ~r1_collsp(sK0,X3,sK1,sK2) | ~r1_collsp(sK0,sK2,X3,sK1)) & r1_collsp(sK0,sK1,sK2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) => ((~r1_collsp(sK0,sK3,sK2,sK1) | ~r1_collsp(sK0,sK1,sK3,sK2) | ~r1_collsp(sK0,sK2,sK1,sK3) | ~r1_collsp(sK0,sK3,sK1,sK2) | ~r1_collsp(sK0,sK2,sK3,sK1)) & r1_collsp(sK0,sK1,sK2,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
28. ((((~r1_collsp(sK0,sK3,sK2,sK1) | ~r1_collsp(sK0,sK1,sK3,sK2) | ~r1_collsp(sK0,sK2,sK1,sK3) | ~r1_collsp(sK0,sK3,sK1,sK2) | ~r1_collsp(sK0,sK2,sK3,sK1)) & r1_collsp(sK0,sK1,sK2,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_collsp(sK0) & v3_collsp(sK0) & v2_collsp(sK0) & ~v2_struct_0(sK0) [skolemisation 21,27,26,25,24]
35. ~v2_struct_0(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
36. v2_collsp(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
37. v3_collsp(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
38. l1_collsp(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
39. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 28] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
40. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 28] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
41. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 28] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
42. r1_collsp(sK0,sK1,sK2,sK3) [cnf transformation 28] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
43. ~r1_collsp(sK0,sK3,sK2,sK1) | ~r1_collsp(sK0,sK1,sK3,sK2) | ~r1_collsp(sK0,sK2,sK1,sK3) | ~r1_collsp(sK0,sK3,sK1,sK2) | ~r1_collsp(sK0,sK2,sK3,sK1) [cnf transformation 28] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
44. ~r1_collsp(X0,X1,X2,X3) | r1_collsp(X0,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0) [cnf transformation 23] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. ~r1_collsp(X0,X1,X2,X3) | r1_collsp(X0,X1,X3,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0) [cnf transformation 23] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. 1 <=> r1_collsp(sK0,sK2,sK3,sK1) [avatar definition]
54. r1_collsp(sK0,sK2,sK3,sK1) <- (1) [avatar component clause 53] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
57. 2 <=> r1_collsp(sK0,sK3,sK1,sK2) [avatar definition]
59. ~r1_collsp(sK0,sK3,sK1,sK2) <- (~2) [avatar component clause 57] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
61. 3 <=> r1_collsp(sK0,sK2,sK1,sK3) [avatar definition]
62. r1_collsp(sK0,sK2,sK1,sK3) <- (3) [avatar component clause 61] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
65. 4 <=> r1_collsp(sK0,sK1,sK3,sK2) [avatar definition]
66. r1_collsp(sK0,sK1,sK3,sK2) <- (4) [avatar component clause 65] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
69. 5 <=> r1_collsp(sK0,sK3,sK2,sK1) [avatar definition]
72. ~1 | ~2 | ~3 | ~4 | ~5 [avatar split clause 43,69,65,61,57,53]
74. 6 <=> r1_collsp(sK0,sK1,sK2,sK3) [avatar definition]
76. r1_collsp(sK0,sK1,sK2,sK3) <- (6) [avatar component clause 74] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
77. 6 [avatar split clause 42,74]
79. 7 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
81. m1_subset_1(sK3,u1_struct_0(sK0)) <- (7) [avatar component clause 79] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
82. 7 [avatar split clause 41,79]
84. 8 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
86. m1_subset_1(sK2,u1_struct_0(sK0)) <- (8) [avatar component clause 84] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
87. 8 [avatar split clause 40,84]
89. 9 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
91. m1_subset_1(sK1,u1_struct_0(sK0)) <- (9) [avatar component clause 89] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
92. 9 [avatar split clause 39,89]
94. 10 <=> l1_collsp(sK0) [avatar definition]
96. l1_collsp(sK0) <- (10) [avatar component clause 94] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
97. 10 [avatar split clause 38,94]
99. 11 <=> v3_collsp(sK0) [avatar definition]
101. v3_collsp(sK0) <- (11) [avatar component clause 99] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
102. 11 [avatar split clause 37,99]
104. 12 <=> v2_collsp(sK0) [avatar definition]
106. v2_collsp(sK0) <- (12) [avatar component clause 104] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
107. 12 [avatar split clause 36,104]
109. 13 <=> v2_struct_0(sK0) [avatar definition]
111. ~v2_struct_0(sK0) <- (~13) [avatar component clause 109] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
112. ~13 [avatar split clause 35,109]
138. r1_collsp(sK0,sK2,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6) [resolution 44,76] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
139. r1_collsp(sK0,sK2,sK1,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7) [subsumption resolution 138,81] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
140. r1_collsp(sK0,sK2,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8) [subsumption resolution 139,86] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
141. r1_collsp(sK0,sK2,sK1,sK3) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9) [subsumption resolution 140,91] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
142. r1_collsp(sK0,sK2,sK1,sK3) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10) [subsumption resolution 141,96] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
143. r1_collsp(sK0,sK2,sK1,sK3) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10, 11) [subsumption resolution 142,101] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
144. r1_collsp(sK0,sK2,sK1,sK3) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10, 11, 12) [subsumption resolution 143,106] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
145. r1_collsp(sK0,sK2,sK1,sK3) <- (6, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 144,111] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
146. 3 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar split clause 145,109,104,99,94,89,84,79,74,61]
149. r1_collsp(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6) [resolution 45,76] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
150. r1_collsp(sK0,sK2,sK3,sK1) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3) [resolution 45,62] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
152. r1_collsp(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7) [subsumption resolution 149,81] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
153. r1_collsp(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8) [subsumption resolution 152,86] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
154. r1_collsp(sK0,sK1,sK3,sK2) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9) [subsumption resolution 153,91] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
155. r1_collsp(sK0,sK1,sK3,sK2) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10) [subsumption resolution 154,96] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
156. r1_collsp(sK0,sK1,sK3,sK2) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10, 11) [subsumption resolution 155,101] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
157. r1_collsp(sK0,sK1,sK3,sK2) | v2_struct_0(sK0) <- (6, 7, 8, 9, 10, 11, 12) [subsumption resolution 156,106] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
158. r1_collsp(sK0,sK1,sK3,sK2) <- (6, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 157,111] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
168. r1_collsp(sK0,sK2,sK3,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7) [subsumption resolution 150,81] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
169. r1_collsp(sK0,sK2,sK3,sK1) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7, 9) [subsumption resolution 168,91] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
170. r1_collsp(sK0,sK2,sK3,sK1) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7, 8, 9) [subsumption resolution 169,86] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
171. r1_collsp(sK0,sK2,sK3,sK1) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7, 8, 9, 10) [subsumption resolution 170,96] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
172. r1_collsp(sK0,sK2,sK3,sK1) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (3, 7, 8, 9, 10, 11) [subsumption resolution 171,101] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
173. r1_collsp(sK0,sK2,sK3,sK1) | v2_struct_0(sK0) <- (3, 7, 8, 9, 10, 11, 12) [subsumption resolution 172,106] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
174. r1_collsp(sK0,sK2,sK3,sK1) <- (3, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 173,111] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
175. 1 | ~3 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar split clause 174,109,104,99,94,89,84,79,61,53]
176. 4 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar split clause 158,109,104,99,94,89,84,79,74,65]
178. r1_collsp(sK0,sK3,sK2,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1) [resolution 54,44] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
179. r1_collsp(sK0,sK3,sK2,sK1) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 9) [subsumption resolution 178,91] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
180. r1_collsp(sK0,sK3,sK2,sK1) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 7, 9) [subsumption resolution 179,81] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
181. r1_collsp(sK0,sK3,sK2,sK1) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 7, 8, 9) [subsumption resolution 180,86] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
182. r1_collsp(sK0,sK3,sK2,sK1) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 7, 8, 9, 10) [subsumption resolution 181,96] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
183. r1_collsp(sK0,sK3,sK2,sK1) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (1, 7, 8, 9, 10, 11) [subsumption resolution 182,101] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
184. r1_collsp(sK0,sK3,sK2,sK1) | v2_struct_0(sK0) <- (1, 7, 8, 9, 10, 11, 12) [subsumption resolution 183,106] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
185. r1_collsp(sK0,sK3,sK2,sK1) <- (1, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 184,111] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
186. 5 | ~1 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar split clause 185,109,104,99,94,89,84,79,53,69]
189. r1_collsp(sK0,sK3,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (4) [resolution 66,44] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
191. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4) [subsumption resolution 189,59] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
192. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 8) [subsumption resolution 191,86] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
193. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 7, 8) [subsumption resolution 192,81] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
194. ~l1_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 7, 8, 9) [subsumption resolution 193,91] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
195. ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 7, 8, 9, 10) [subsumption resolution 194,96] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
196. ~v2_collsp(sK0) | v2_struct_0(sK0) <- (~2, 4, 7, 8, 9, 10, 11) [subsumption resolution 195,101] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
197. v2_struct_0(sK0) <- (~2, 4, 7, 8, 9, 10, 11, 12) [subsumption resolution 196,106] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
198. $false <- (~2, 4, 7, 8, 9, 10, 11, 12, ~13) [subsumption resolution 197,111] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
199. 2 | ~4 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | 13 [avatar contradiction clause 198]
200. $false [avatar sat refutation 72,77,82,87,92,97,102,107,112,146,175,176,186,199] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:237,thDist:-237}
% SZS output end Proof for hessenbe__t1_hessenbe
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 92
% Active clauses: 21
% Passive clauses: 22
% Final active clauses: 20
% Final passive clauses: 1
% Main loop iterations started: 21
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 44
% 
% >>> Deletion Inferences
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 8
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 5
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 13
% 
% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

