i: [1,197,0,2,1,1,0,0] l27_xreal_0
a: [3,314,0,2,1,197]
s: 314
i: [1,196,0,2,1,1,0,0] l27_xreal_0
a: [3,319,0,2,1,196]
s: 319
i: [1,199,0,3,1,1,0,0] l27_xreal_0
a: [3,304,0,3,1,199]
s: 304
i: [1,198,0,3,1,1,0,0] l27_xreal_0
a: [3,309,0,3,1,198]
s: 309
i: [1,200,0,5,1,1,0,0] l27_xreal_0
a: [3,299,0,5,1,200]
s: 299
i: [1,201,0,2,1,0,0,255] fc1_numbers
a: [3,324,0,2,1,201]
s: 324
i: [1,202,0,2,1,0,0,255] spc1_boole
a: [3,329,0,2,1,202]
s: 329
i: [1,203,0,2,1,0,0,255] fc1_xboole_0
a: [3,334,0,2,1,203]
s: 334
i: [1,212,0,2,1,0,0,255] fc6_ordinal1
a: [3,374,0,2,1,212]
s: 374
i: [1,211,0,2,1,0,0,255] fc6_ordinal1
a: [3,379,0,2,1,211]
s: 379
i: [1,213,0,2,1,0,0,255] spc1_numerals
a: [3,399,0,2,1,213]
s: 399
i: [1,275,0,2,1,0,0,2] rc2_xboole_0
a: [3,411,0,2,1,275]
s: 411
i: [1,276,0,2,1,0,0,255] rc4_ordinal1
a: [3,416,0,2,1,276]
s: 416
i: [1,277,0,2,1,0,0,1] rc1_xreal_0
a: [3,421,0,2,1,277]
s: 421
i: [1,278,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,426,0,2,1,278]
s: 426
i: [1,279,0,2,1,0,0,255] rc1_ordinal1
a: [3,431,0,2,1,279]
s: 431
i: [1,280,0,2,1,0,0,2] rc1_xxreal_0
a: [3,436,0,2,1,280]
s: 436
i: [1,281,0,2,1,0,0,2] rc1_xboole_0
a: [3,441,0,2,1,281]
s: 441
i: [1,285,0,2,1,0,0,255] rc2_ordinal1
a: [3,446,0,2,1,285]
s: 446
i: [1,284,0,2,1,0,0,255] rc2_ordinal1
a: [3,451,0,2,1,284]
s: 451
i: [1,283,0,2,1,0,0,255] rc2_ordinal1
a: [3,456,0,2,1,283]
s: 456
i: [1,282,0,2,1,0,0,255] rc2_ordinal1
a: [3,461,0,2,1,282]
s: 461
i: [1,287,0,2,1,0,0,255] rc5_ordinal1
a: [3,466,0,2,1,287]
s: 466
i: [1,286,0,2,1,0,0,255] rc5_ordinal1
a: [3,471,0,2,1,286]
s: 471
i: [1,289,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,476,0,2,1,289]
s: 476
i: [1,288,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,481,0,2,1,288]
s: 481
i: [1,291,0,2,1,0,0,255] rc3_xxreal_0
a: [3,486,0,2,1,291]
s: 486
i: [1,290,0,2,1,0,0,255] rc3_xxreal_0
a: [3,491,0,2,1,290]
s: 491
i: [1,293,0,2,1,0,0,255] rc2_xxreal_0
a: [3,496,0,2,1,293]
s: 496
i: [1,292,0,2,1,0,0,255] rc2_xxreal_0
a: [3,501,0,2,1,292]
s: 501
i: [1,295,0,2,1,0,0,2] rc4_xxreal_0
a: [3,506,0,2,1,295]
s: 506
i: [1,294,0,2,1,0,0,2] rc4_xxreal_0
a: [3,511,0,2,1,294]
s: 511
i: [1,262,0,3,1,0,0,255] reflexivity_r1_tarski
s: 262
i: [1,204,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,339,0,3,1,204]
s: 339
i: [1,205,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,344,0,3,1,205]
s: 344
r: 334
d: [2,513,0,2,1,2,23,334,344]
a: [3,517,0,2,1,513]
s: 517
i: [1,206,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,349,0,3,1,206]
s: 349
i: [1,207,0,3,1,0,0,255] t1_numerals
a: [3,354,0,3,1,207]
s: 354
i: [1,216,0,3,1,0,0,255] spc1_numerals
a: [3,384,0,3,1,216]
s: 384
i: [1,215,0,3,1,0,0,255] spc1_numerals
a: [3,389,0,3,1,215]
s: 389
d: [2,519,0,3,1,2,22,354,344]
a: [3,523,0,3,1,519]
s: 523
d: [2,525,0,3,1,2,22,389,339]
a: [3,529,0,3,1,525]
s: 529
i: [1,218,0,4,2,0,0,255] cc2_xreal_0
s: 218
d: [2,531,1,2,1,1,40,218,416]
a: [3,536,1,2,1,531]
s: 536
i: [1,219,0,4,2,0,0,255] cc2_xcmplx_0
s: 219
i: [1,220,0,4,2,0,0,255] cc6_ordinal1
s: 220
d: [2,532,1,2,1,1,40,218,466]
a: [3,541,1,2,1,532]
s: 541
i: [1,221,0,4,2,0,0,255] cc2_xxreal_0
s: 221
d: [2,543,1,2,1,1,40,219,416]
a: [3,548,1,2,1,543]
s: 548
i: [1,222,0,4,2,0,0,1] cc3_xreal_0
s: 222
i: [1,223,0,4,2,0,0,1] cc4_xreal_0
s: 223
d: [2,579,1,2,1,1,40,222,319]
a: [3,587,1,2,1,579]
s: 587
i: [1,228,0,4,2,0,0,255] cc1_ordinal1
s: 228
d: [2,580,1,2,1,1,40,222,314]
a: [3,592,1,2,1,580]
s: 592
i: [1,229,0,4,2,0,0,255] cc1_ordinal1
s: 229
i: [1,231,0,4,2,0,0,255] cc9_ordinal1
s: 231
d: [2,600,1,2,1,1,40,223,319]
a: [3,608,1,2,1,600]
s: 608
i: [1,232,0,4,2,0,0,255] cc7_ordinal1
s: 232
d: [2,601,1,2,1,1,40,223,314]
a: [3,613,1,2,1,601]
s: 613
i: [1,233,0,4,2,0,0,255] cc3_ordinal1
s: 233
i: [1,261,0,4,1,0,0,255] existence_m1_subset_1
s: 261
d: [2,544,1,2,1,1,40,219,466]
a: [3,553,1,2,1,544]
s: 553
i: [1,208,0,4,1,0,0,255] dt_k5_numbers
a: [3,359,0,4,1,208]
s: 359
d: [2,555,1,2,1,1,40,220,416]
a: [3,560,1,2,1,555]
s: 560
i: [1,209,0,4,1,0,0,255] dt_k6_numbers
a: [3,364,0,4,1,209]
s: 364
i: [1,214,0,4,1,0,0,255] spc1_numerals
a: [3,394,0,4,1,214]
s: 394
d: [2,556,1,2,1,1,40,220,466]
a: [3,565,1,2,1,556]
s: 565
d: [2,702,0,4,1,2,22,359,339]
a: [3,706,0,4,1,702]
s: 706
d: [2,567,1,2,1,1,40,221,416]
a: [3,572,1,2,1,567]
s: 572
d: [2,720,0,4,1,2,22,364,339]
a: [3,724,0,4,1,720]
s: 724
d: [2,726,0,4,1,2,22,394,339]
a: [3,730,0,4,1,726]
s: 730
d: [2,568,1,2,1,1,40,221,466]
a: [3,577,1,2,1,568]
s: 577
i: [1,234,0,5,2,0,0,255] t6_boole
s: 234
d: [2,582,1,2,1,1,40,222,421]
a: [3,597,1,2,1,582]
s: 597
i: [1,235,0,5,2,0,0,255] cc8_ordinal1
s: 235
i: [1,236,0,5,2,0,0,255] cc1_xreal_0
s: 236
d: [2,603,1,2,1,1,40,223,421]
a: [3,619,1,2,1,603]
s: 619
i: [1,237,0,5,2,0,0,255] cc1_xcmplx_0
s: 237
d: [2,622,1,2,1,1,40,228,374]
a: [3,628,1,2,1,622]
s: 628
i: [1,258,0,5,2,0,0,255] t7_boole
s: 258
i: [1,210,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,369,0,5,1,210]
s: 369
d: [2,623,1,2,1,1,40,228,431]
a: [3,633,1,2,1,623]
s: 633
i: [1,272,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,406,0,5,2,272]
s: 406
d: [2,635,1,2,1,1,40,229,374]
a: [3,641,1,2,1,635]
s: 641
d: [2,744,0,5,2,1,22,234,344]
s: 744
i: [1,242,0,6,3,0,0,255] cc2_ordinal1
s: 242
d: [2,785,1,5,2,2,44,304,744]
s: 785
i: [1,245,0,6,3,0,0,255] cc5_xxreal_0
s: 245
d: [2,786,1,5,2,2,44,309,744]
s: 786
i: [1,247,0,6,3,0,0,255] cc5_xxreal_0
s: 247
i: [1,248,0,6,3,0,0,255] cc3_xxreal_0
s: 248
d: [2,784,1,7,2,2,44,299,744]
s: 784
i: [1,250,0,6,3,0,0,255] cc3_xxreal_0
s: 250
d: [2,814,2,8,2,3,40,784,406]
a: [3,823,2,4,1,814]
s: 823
i: [1,252,0,6,3,0,0,255] cc7_xxreal_0
s: 252
i: [1,253,0,6,3,0,0,255] cc7_xxreal_0
s: 253
d: [2,636,1,2,1,1,40,229,431]
a: [3,646,1,2,1,636]
s: 646
i: [1,254,0,6,2,0,0,255] t1_subset
s: 254
d: [2,648,1,2,1,1,40,231,511]
a: [3,654,1,2,1,648]
s: 654
i: [1,263,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 263
i: [1,226,0,7,2,0,0,1] t2_arithm
s: 226
d: [2,649,1,2,1,1,40,231,441]
a: [3,659,1,2,1,649]
s: 659
i: [1,227,0,7,2,0,0,255] t3_arithm
s: 227
d: [2,650,1,2,1,1,40,231,517]
a: [3,664,1,2,1,650]
s: 664
i: [1,230,0,7,3,0,0,255] cc5_ordinal1
s: 230
i: [1,256,0,7,2,0,0,255] t3_subset
s: 256
d: [2,666,1,2,1,1,40,232,511]
a: [3,672,1,2,1,666]
s: 672
i: [1,257,0,7,3,0,0,2] t8_boole
s: 257
d: [2,667,1,2,1,1,40,232,441]
a: [3,677,1,2,1,667]
s: 677
i: [1,217,0,8,3,0,0,255] cc10_ordinal1
s: 217
i: [1,239,0,8,4,0,0,255] cc4_xxreal_0
s: 239
d: [2,815,2,6,2,3,40,784,785]
a: [3,827,2,2,1,815]
s: 827
r: 827
i: [1,241,0,8,4,0,0,255] cc6_xxreal_0
s: 241
a: [3,828,2,2,1,815]
s: 828
i: [1,243,0,8,4,0,0,255] cc8_xxreal_0
s: 243
i: [1,255,0,8,3,0,0,255] t2_subset
s: 255
d: [2,899,2,6,3,1,40,239,608]
a: [3,911,2,2,1,899]
s: 911
r: 911
i: [1,269,0,8,3,0,0,1] fc6_xreal_0
s: 269
a: [3,910,2,2,1,899]
s: 910
i: [1,270,0,8,3,0,0,1] fc3_xcmplx_0
s: 270
i: [1,260,0,9,3,0,0,255] t5_subset
s: 260
a: [3,915,2,2,1,899]
s: 915
i: [1,259,0,10,3,0,0,255] t4_subset
s: 259
d: [2,900,2,6,3,1,40,239,613]
a: [3,924,2,2,1,900]
s: 924
i: [1,273,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 273
i: [1,271,0,11,3,0,0,1] commutativity_k3_xcmplx_0
s: 271
a: [3,819,2,4,1,814]
s: 819
i: [1,268,0,12,5,0,0,1] fc8_xcmplx_0
s: 268
d: [2,668,1,2,1,1,40,232,517]
a: [3,682,1,2,1,668]
s: 682
i: [1,224,0,13,5,0,0,1] l5_xreal_0
s: 224
i: [1,267,0,14,4,0,0,255] existence_m2_subset_1
s: 267
d: [2,1068,1,11,3,1,40,273,299]
a: [3,1074,1,2,1,1068]
s: 1074
i: [1,225,0,15,5,0,0,1] l25_xreal_0
s: 225
d: [2,1151,1,10,4,1,40,224,304]
d: [2,1158,1,8,3,2,21,1151,314]
a: [3,1162,1,3,1,1158]
s: 1162
i: [1,264,0,15,5,0,0,255] dt_m2_subset_1
s: 264
i: [1,265,0,15,5,0,0,255] redefinition_m2_subset_1
s: 265
d: [2,1153,1,10,4,1,40,224,309]
d: [2,1168,1,8,3,2,21,1153,319]
a: [3,1172,1,3,1,1168]
s: 1172
i: [1,266,0,15,5,0,0,255] redefinition_m2_subset_1
s: 266
d: [2,1180,2,4,2,2,44,1074,744]
s: 1180
i: [1,274,0,17,4,0,0,1] spc7_arithm
s: 274
d: [2,684,1,2,1,1,40,233,511]
a: [3,690,1,2,1,684]
s: 690
d: [2,1195,2,5,2,2,44,1162,744]
s: 1195
d: [2,685,1,2,1,1,40,233,441]
a: [3,695,1,2,1,685]
s: 695
d: [2,1213,2,5,2,2,44,1172,744]
s: 1213
d: [2,686,1,2,1,1,40,233,517]
a: [3,700,1,2,1,686]
s: 700
d: [2,746,1,2,1,1,40,235,529]
a: [3,752,1,2,1,746]
s: 752
d: [2,759,1,2,1,1,40,236,384]
a: [3,764,1,2,1,759]
s: 764
d: [2,771,1,2,1,1,40,237,384]
a: [3,776,1,2,1,771]
s: 776
d: [2,792,1,4,2,1,40,245,486]
d: [2,793,1,2,1,2,21,792,491]
a: [3,797,1,2,1,793]
s: 797
d: [2,799,1,4,2,1,40,247,486]
d: [2,800,1,2,1,2,21,799,491]
a: [3,804,1,2,1,800]
s: 804
d: [2,807,1,4,2,1,40,248,496]
d: [2,808,1,2,1,2,21,807,501]
a: [3,812,1,2,1,808]
s: 812
d: [2,835,1,4,2,1,40,230,384]
a: [3,847,1,2,1,835]
s: 847
d: [2,888,1,4,2,1,40,217,706]
a: [3,893,1,2,1,888]
s: 893
d: [2,903,1,6,3,1,40,239,436]
a: [3,959,1,2,1,903]
s: 959
r: 959
a: [3,958,1,2,1,903]
s: 958
a: [3,963,1,2,1,903]
s: 963
d: [2,1086,1,11,4,0,44,271,226]
d: [2,1105,1,9,3,0,18,1086]
a: [3,1113,1,2,1,1105]
s: 1113
d: [2,747,1,3,1,0,40,235,261]
a: [3,757,1,3,1,747]
s: 757
d: [2,708,2,2,1,1,40,560,229]
a: [3,713,2,2,1,708]
s: 713
d: [2,760,1,3,1,0,40,236,261]
a: [3,769,1,3,1,760]
s: 769
d: [2,772,1,3,1,0,40,237,261]
a: [3,781,1,3,1,772]
s: 781
d: [2,709,2,2,1,1,40,560,228]
a: [3,718,2,2,1,709]
s: 718
d: [2,998,1,5,2,1,40,255,523]
d: [2,1003,1,3,1,2,21,998,379]
a: [3,1007,1,3,1,1003]
s: 1007
d: [2,732,2,2,1,1,40,565,229]
a: [3,737,2,2,1,732]
s: 737
d: [2,1000,1,5,2,1,40,255,384]
d: [2,1018,1,3,1,2,21,1000,324]
a: [3,1022,1,3,1,1018]
s: 1022
d: [2,1001,1,5,2,1,40,255,529]
d: [2,1024,1,3,1,2,21,1001,379]
a: [3,1028,1,3,1,1024]
s: 1028
d: [2,733,2,2,1,1,40,565,228]
a: [3,742,2,2,1,733]
s: 742
d: [2,1196,1,11,4,1,40,264,724]
d: [2,1200,1,7,3,2,21,1196,706]
d: [2,1201,1,5,2,3,21,1200,379]
d: [2,1202,1,3,1,4,21,1201,324]
a: [3,1206,1,3,1,1202]
s: 1206
d: [2,854,2,2,1,1,40,672,219]
a: [3,860,2,2,1,854]
s: 860
d: [2,999,1,7,2,1,40,255,706]
a: [3,1012,1,4,1,999]
s: 1012
a: [3,1015,1,3,1,999]
s: 1015
d: [2,855,2,2,1,1,40,672,218]
a: [3,865,2,2,1,855]
s: 865
d: [2,1181,1,21,7,0,40,225,273]
d: [2,1189,1,19,6,1,21,1181,1074]
s: 1189
d: [2,1458,2,15,6,2,44,1189,369]
d: [2,1466,2,10,4,2,18,1458]
d: [2,1467,2,8,3,3,21,1466,764]
a: [3,1475,2,3,1,1467]
s: 1475
d: [2,788,1,5,2,2,44,523,744]
s: 788
d: [2,837,1,5,2,0,40,230,261]
s: 837
a: [3,1470,2,3,1,1467]
s: 1470
d: [2,789,1,6,2,2,44,724,744]
s: 789
d: [2,1490,3,5,2,2,44,1475,744]
s: 1490
d: [2,889,1,6,2,0,40,217,261]
s: 889
d: [2,1002,1,6,2,0,40,255,261]
s: 1002
d: [2,1504,3,5,2,2,44,1470,744]
s: 1504
d: [2,1065,1,6,2,1,40,259,706]
s: 1065
d: [2,867,2,2,1,1,40,677,221]
a: [3,874,2,2,1,867]
s: 874
d: [2,783,1,7,3,1,44,744,744]
s: 783
d: [2,887,1,7,3,0,40,217,256]
s: 887
d: [2,869,2,2,1,1,40,677,219]
a: [3,880,2,2,1,869]
s: 880
d: [2,1062,1,7,2,0,40,260,261]
s: 1062
d: [2,870,2,2,1,1,40,677,218]
a: [3,885,2,2,1,870]
s: 885
d: [2,1085,1,11,4,0,44,271,227]
d: [2,1106,1,9,3,0,18,1085]
a: [3,1109,1,7,2,1106]
s: 1109
a: [3,1117,1,7,2,1105]
s: 1117
d: [2,1546,2,17,7,2,44,1189,1109]
d: [2,1548,2,14,6,2,18,1546]
d: [2,1562,2,12,5,3,21,1548,764]
d: [2,1563,2,9,4,4,21,1562,1470]
d: [2,1564,2,7,3,4,21,1563,222]
d: [2,1565,2,5,2,4,21,1564,223]
s: 1565
r: 225
r: 1189
d: [2,1588,3,0,0,6,52,314,1162,1565]
e: 1588
d: [2,1592,3,2,1,5,40,1565,1162]
d: [2,1598,3,0,0,6,21,1592,314]
e: 1598
d: [2,1593,3,4,2,6,40,1565,1195]
d: [2,1600,3,2,1,7,21,1593,314]
d: [2,1601,3,0,0,8,21,1600,517]
e: 1601
r: 828
r: 1195
r: 1162
r: 924
r: 1565
r: 1504
r: 1470
r: 1490
r: 1475
s: 225
s: 827
a: [3,1166,1,3,1,1158]
s: 1166
r: 299
r: 304
r: 314
r: 592
r: 613
r: 784
r: 785
r: 819
r: 823
r: 827
r: 1166
s: 314
s: 592
s: 304
s: 613
s: 299
s: 1162
s: 827
a: [3,1165,1,3,1,1158]
s: 1165
d: [2,1709,3,0,0,3,52,827,1165,744]
e: 1709
d: [2,1710,3,0,0,3,52,517,827,1165,257]
e: 1710
d: [2,1711,3,0,0,3,52,827,517,1165,257]
e: 1711
r: 827
s: 828
s: 785
d: [2,1691,1,5,2,2,44,304,744]
s: 1691
s: 924
s: 784
d: [2,979,3,2,1,1,40,827,233]
a: [3,985,3,2,1,979]
s: 985
r: 985
d: [2,1697,1,7,2,2,44,299,744]
s: 1697
s: 1189
a: [3,984,3,2,1,979]
s: 984
d: [2,1628,1,21,7,0,40,225,273]
d: [2,1639,1,19,6,1,21,1628,1074]
s: 1639
d: [2,980,3,2,1,1,40,827,232]
a: [3,990,3,2,1,980]
s: 990
d: [2,1771,4,0,0,2,52,984,990,220]
e: 1771
d: [2,1773,4,2,1,1,40,990,220]
d: [2,1777,4,0,0,2,21,1773,984]
e: 1777
r: 990
d: [2,1215,1,11,4,1,40,266,706]
d: [2,1217,1,9,3,2,21,1215,379]
d: [2,1218,1,7,2,3,21,1217,324]
s: 1218
d: [2,851,1,8,3,0,40,256,230]
s: 851
d: [2,981,3,2,1,1,40,827,231]
a: [3,995,3,2,1,981]
s: 995
d: [2,1048,1,8,3,0,40,269,223]
s: 1048
a: [3,1471,2,3,1,1467]
s: 1471
d: [2,1049,1,8,3,0,40,269,222]
s: 1049
d: [2,1060,1,8,3,0,40,260,256]
s: 1060
a: [3,1474,2,3,1,1467]
s: 1474
d: [2,1066,1,8,2,0,40,259,261]
s: 1066
s: 819
d: [2,1814,3,0,0,3,52,314,319,819,1048]
e: 1814
d: [2,1815,3,4,2,1,40,819,1048]
d: [2,1817,3,2,1,2,21,1815,319]
d: [2,1818,3,0,0,3,21,1817,314]
e: 1818
r: 819
a: [3,1079,1,5,1,1068]
s: 1079
d: [2,831,1,9,3,1,44,226,744]
s: 831
a: [3,818,2,4,1,814]
s: 818
r: 406
d: [2,1656,0,3,1,2,23,304,1166]
a: [3,1673,0,3,1,1656]
s: 1673
s: 823
a: [3,403,0,2,1,272]
s: 403
r: 1189
r: 273
r: 248
r: 247
r: 245
r: 239
d: [2,1921,1,0,0,2,40,403,506]
e: 1921
d: [2,1920,1,0,0,2,40,403,501]
e: 1920
d: [2,1919,1,0,0,2,40,403,491]
e: 1919
d: [2,1918,2,0,0,2,40,403,577]
e: 1918
d: [2,1917,3,0,0,2,40,403,874]
e: 1917
d: [2,1916,1,0,0,2,40,403,436]
e: 1916
d: [2,1915,2,0,0,2,40,403,619]
e: 1915
d: [2,1914,2,0,0,2,40,403,572]
e: 1914
d: [2,1913,2,0,0,2,40,403,613]
e: 1913
d: [2,1912,2,0,0,2,40,403,608]
e: 1912
d: [2,1911,3,0,0,2,40,403,818]
e: 1911
d: [2,1909,2,0,0,2,40,403,1074]
e: 1909
d: [2,1896,2,0,0,2,52,1074,403]
e: 1896
d: [2,1897,3,0,0,4,52,517,403,1180]
e: 1897
d: [2,1898,2,0,0,2,52,764,403,223]
e: 1898
d: [2,1899,2,0,0,2,52,682,403,221]
e: 1899
d: [2,1900,2,0,0,2,52,764,764,403,1048]
e: 1900
d: [2,1901,2,0,0,3,52,764,319,403,1048]
e: 1901
d: [2,1902,2,0,0,3,52,764,314,403,1048]
e: 1902
d: [2,1903,2,0,0,3,52,764,769,403,1048]
e: 1903
d: [2,1904,2,0,0,3,52,764,536,403,1048]
e: 1904
d: [2,1905,2,0,0,3,52,764,421,403,1048]
e: 1905
d: [2,1906,3,0,0,3,52,764,885,403,1048]
e: 1906
d: [2,1907,2,0,0,3,52,764,541,403,1048]
e: 1907
d: [2,1908,3,0,0,3,52,764,865,403,1048]
e: 1908
r: 403
d: [2,1947,0,4,2,2,21,245,1180]
s: 1947
d: [2,1948,0,4,2,2,21,248,1180]
s: 1948
s: 406
d: [2,1864,3,12,3,1,40,818,239]
a: [3,1869,3,4,1,1864]
s: 1869
s: 247
s: 239
d: [2,1892,1,5,2,2,44,1673,744]
s: 1892
s: 273
d: [2,1616,2,10,4,3,21,1562,222]
d: [2,1617,2,8,3,3,21,1616,223]
a: [3,1620,2,5,2,1617]
s: 1620
r: 225
d: [2,2001,3,0,0,3,52,764,1474,1620]
e: 2001
d: [2,2002,3,0,0,3,52,314,1162,1620]
e: 2002
d: [2,2006,3,2,1,2,40,1620,1474]
d: [2,2013,3,0,0,3,21,2006,764]
e: 2013
d: [2,2007,3,2,1,2,40,1620,1162]
d: [2,2015,3,0,0,3,21,2007,314]
e: 2015
r: 1620
s: 225
d: [2,1949,1,15,5,1,21,1189,1048]
s: 1949
d: [2,1701,2,5,2,2,44,1162,744]
s: 1701
d: [2,1989,1,21,7,0,40,273,225]
d: [2,1996,1,19,6,1,21,1989,1074]
d: [2,1997,1,15,5,1,21,1996,1048]
s: 1997
d: [2,1797,3,5,2,2,44,1471,744]
s: 1797
d: [2,2018,1,21,7,0,40,225,273]
d: [2,2030,1,19,6,1,21,2018,1074]
d: [2,2031,1,15,5,1,21,2030,1048]
s: 2031
d: [2,1064,1,9,3,0,40,259,256]
s: 1064
d: [2,1812,3,5,2,2,44,1474,744]
s: 1812
d: [2,997,1,10,3,0,40,255,256]
s: 997
d: [2,1719,2,13,4,2,40,784,273]
d: [2,1723,2,11,3,3,21,1719,1180]
a: [3,1726,2,7,2,1723]
s: 1726
d: [2,1223,1,15,4,1,44,274,369]
d: [2,1250,1,13,3,1,18,1223]
a: [3,1253,1,11,2,1250]
s: 1253
d: [2,1229,1,17,5,0,44,274,226]
d: [2,1264,1,13,4,0,21,1229,270]
a: [3,1267,1,11,3,1264]
s: 1267
r: 828
r: 924
r: 984
r: 1701
r: 1162
r: 1165
r: 1797
r: 1471
r: 1812
r: 1474
s: 828
d: [2,1655,0,5,1,2,23,299,1166]
a: [3,1668,0,5,1,1655]
s: 1668
d: [2,2223,1,5,2,1,44,1668,226]
d: [2,2234,1,2,1,2,21,2223,1673]
d: [2,2235,1,0,0,3,21,2234,587]
e: 2235
d: [2,2224,2,7,3,2,44,1668,831]
d: [2,2237,2,4,2,3,21,2224,1673]
d: [2,2238,2,2,1,4,21,2237,587]
d: [2,2239,2,0,0,5,21,2238,517]
e: 2239
r: 1668
a: [3,923,2,2,1,900]
s: 923
a: [3,1667,0,5,1,1655]
s: 1667
s: 1162
a: [3,928,2,2,1,900]
s: 928
s: 1165
s: 985
d: [2,2222,1,11,3,1,40,1668,273]
d: [2,2228,1,9,2,2,21,2222,1074]
a: [3,2231,1,5,1,2228]
s: 2231
d: [2,2281,2,5,2,1,44,2231,226]
d: [2,2285,2,2,1,2,21,2281,1673]
d: [2,2286,2,0,0,3,21,2285,587]
e: 2286
d: [2,2282,2,7,3,2,44,2231,831]
d: [2,2288,2,4,2,3,21,2282,1673]
d: [2,2289,2,2,1,4,21,2288,587]
d: [2,2290,2,0,0,5,21,2289,517]
e: 2290
r: 2231
d: [2,2245,1,16,4,1,40,1667,224]
a: [3,2254,1,4,1,2245]
s: 2254
r: 2254
s: 990
a: [3,2253,1,4,1,2245]
s: 2253
d: [2,1728,4,2,1,1,40,985,229]
a: [3,1733,4,2,1,1728]
s: 1733
a: [3,2232,1,5,1,2228]
s: 2232
a: [3,2258,1,5,1,2245]
s: 2258
r: 2232
r: 2253
r: 1667
d: [2,1729,4,2,1,1,40,985,228]
a: [3,1738,4,2,1,1729]
s: 1738
d: [2,2248,1,9,3,1,44,1667,271]
d: [2,2260,1,7,2,2,21,2248,587]
d: [2,2261,1,5,1,3,21,2260,1113]
a: [3,2265,1,5,1,2261]
s: 2265
s: 1471
d: [2,2250,1,7,2,2,44,1667,744]
s: 2250
d: [2,1236,1,18,6,0,44,270,274]
d: [2,1241,1,16,5,0,18,1236]
d: [2,1276,1,12,4,0,21,1241,270]
s: 1276
s: 1474
d: [2,1198,1,21,7,0,40,264,267]
d: [2,1199,1,13,4,0,18,1198]
s: 1199
d: [2,1489,3,10,4,1,40,1475,224]
d: [2,1491,3,8,3,2,21,1489,764]
a: [3,1494,3,3,1,1491]
s: 1494
d: [2,1210,1,21,7,0,40,265,267]
d: [2,1211,1,13,4,0,18,1210]
s: 1211
d: [2,1224,1,17,5,0,44,274,227]
d: [2,1249,1,15,4,0,18,1224]
a: [3,1257,1,13,3,1249]
s: 1257
d: [2,1662,2,4,1,2,23,819,1166]
a: [3,1678,2,4,1,1662]
s: 1678
d: [2,1225,1,17,5,0,44,274,226]
d: [2,1248,1,15,4,0,18,1225]
a: [3,1261,1,13,3,1248]
s: 1261
d: [2,1663,2,4,1,2,23,823,1166]
a: [3,1683,2,4,1,1663]
s: 1683
d: [2,1214,1,14,5,0,40,266,256]
s: 1214
d: [2,1186,1,19,7,0,44,225,271]
d: [2,1190,1,17,6,0,21,1186,222]
d: [2,1191,1,15,5,0,21,1190,222]
s: 1191
d: [2,2602,2,0,0,5,52,314,319,1162,1172,1079,1191]
e: 2602
d: [2,2603,2,10,4,1,40,1191,1079]
d: [2,2623,2,7,3,2,21,2603,1172]
d: [2,2624,2,4,2,3,21,2623,1162]
d: [2,2625,2,2,1,4,21,2624,319]
d: [2,2626,2,0,0,5,21,2625,314]
e: 2626
d: [2,2604,3,12,5,1,40,1191,1726]
d: [2,2628,3,9,4,2,21,2604,1172]
d: [2,2629,3,6,3,3,21,2628,1162]
d: [2,2630,3,4,2,4,21,2629,319]
d: [2,2631,3,2,1,5,21,2630,314]
d: [2,2632,3,0,0,6,21,2631,517]
e: 2632
r: 910
r: 1213
r: 1172
r: 915
s: 911
a: [3,1171,1,3,1,1168]
s: 1171
a: [3,914,2,2,1,899]
s: 914
a: [3,1176,1,3,1,1168]
s: 1176
r: 299
r: 309
r: 319
r: 587
r: 608
r: 784
r: 786
r: 818
r: 823
r: 911
r: 914
r: 1079
r: 1171
r: 1726
r: 1869
r: 2250
r: 2258
r: 2265
d: [2,2679,0,5,1,2,23,299,1176]
a: [3,2700,0,5,1,2679]
s: 2700
d: [2,2756,2,5,2,2,44,2700,1117]
d: [2,2761,2,2,1,3,21,2756,1673]
d: [2,2762,2,0,0,4,21,2761,592]
e: 2762
d: [2,745,1,2,1,1,40,235,523]
d: [2,1094,1,11,4,0,44,226,271]
d: [2,1098,1,9,3,0,18,1094]
d: [2,1135,2,2,1,1,40,682,221]
d: [2,1137,2,2,1,1,40,682,219]
d: [2,1138,2,2,1,1,40,682,218]
d: [2,1687,1,2,1,1,40,314,222]
d: [2,1690,1,10,4,1,40,304,224]
d: [2,1692,1,8,3,2,21,1690,314]
d: [2,1718,2,8,2,3,40,784,406]
d: [2,1721,2,4,1,4,21,1718,1180]
d: [2,1460,2,17,7,1,44,1189,226]
d: [2,1465,2,14,6,1,18,1460]
d: [2,1887,2,12,5,1,21,1465,1074]
d: [2,1888,2,10,4,1,21,1887,222]
d: [2,1070,1,7,3,0,42,273]
d: [2,1071,1,5,2,0,18,1070]
d: [2,1991,2,13,4,2,40,273,784]
d: [2,1998,2,9,3,3,21,1991,818]
d: [2,1999,2,7,2,4,21,1998,1180]
d: [2,1894,3,6,3,2,44,823,831]
d: [2,1895,3,4,2,2,18,1894]
d: [2,2192,3,2,1,3,21,1895,587]
d: [2,2645,3,9,4,2,21,2604,1162]
d: [2,2646,3,7,3,3,21,2645,319]
d: [2,2647,3,5,2,4,21,2646,314]
d: [2,2648,3,3,1,5,21,2647,517]
d: [2,2674,2,10,4,1,40,1171,224]
d: [2,2676,2,7,3,2,21,2674,309]
d: [2,2677,2,5,2,3,21,2676,319]
f: 2762,2679,2677,2648,2192,2015,1999,1071,1888,1818,1721,1692,1687,1138,1137,1135,1098,745,579,519,513,207,205,203,196,197,198,199,200
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_0__l27_xreal_0
% SZS output start Proof for xreal_0__l27_xreal_0
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(X0,k6_numbers)) => r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers)))) [input l27_xreal_0]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(X0,k6_numbers)) => r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers)))) [negated conjecture 1]
9. ! [X0] : (v7_ordinal1(X0) => v1_xcmplx_0(X0)) [input cc2_xcmplx_0]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
13. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
27. ! [X0,X1] : ((v1_xcmplx_0(X1) & v1_xcmplx_0(X0)) => k3_xcmplx_0(X0,X1) = k3_xcmplx_0(X1,X0)) [input commutativity_k3_xcmplx_0]
28. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
41. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
44. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k3_xcmplx_0(X0,X1))) [input fc6_xreal_0]
47. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X1)) & ~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(X0,k6_numbers)))) [input l25_xreal_0]
48. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input l5_xreal_0]
63. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
72. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
74. ! [X0] : (v1_xcmplx_0(X0) => k6_numbers = k3_xcmplx_0(X0,k6_numbers)) [input t2_arithm]
80. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
94. ? [X0] : (? [X1] : ((~r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & (r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(X0,k6_numbers))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
95. ? [X0] : (? [X1] : (~r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 94]
97. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
98. ! [X0] : (v1_xcmplx_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 9]
100. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
101. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 13]
102. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
103. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 48]
104. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 103]
105. ! [X0] : (! [X1] : ((~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X1)) | r1_xxreal_0(k6_numbers,X1) | r1_xxreal_0(X0,k6_numbers)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 47]
106. ! [X0] : (! [X1] : (~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X1)) | r1_xxreal_0(k6_numbers,X1) | r1_xxreal_0(X0,k6_numbers) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 105]
107. ! [X0] : (k6_numbers = k3_xcmplx_0(X0,k6_numbers) | ~v1_xcmplx_0(X0)) [ennf transformation 74]
114. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 80]
115. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
150. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 44]
151. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 150]
154. ! [X0,X1] : (k3_xcmplx_0(X0,X1) = k3_xcmplx_0(X1,X0) | (~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0))) [ennf transformation 27]
155. ! [X0,X1] : (k3_xcmplx_0(X0,X1) = k3_xcmplx_0(X1,X0) | ~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0)) [flattening 154]
158. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 28]
159. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 158]
162. ? [X0] : (? [X1] : (~r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~r1_xxreal_0(k3_xcmplx_0(sK0,X1),k6_numbers) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
163. ? [X1] : (~r1_xxreal_0(k3_xcmplx_0(sK0,X1),k6_numbers) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) => (~r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) & r1_xxreal_0(k6_numbers,sK1) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) [choice axiom]
164. (~r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) & r1_xxreal_0(k6_numbers,sK1) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 95,163,162]
196. v1_xreal_0(sK0) [cnf transformation 164] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
197. v1_xreal_0(sK1) [cnf transformation 164] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
198. r1_xxreal_0(sK0,k6_numbers) [cnf transformation 164] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
199. r1_xxreal_0(k6_numbers,sK1) [cnf transformation 164] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
200. ~r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) [cnf transformation 164] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
203. v1_xboole_0(k1_xboole_0) [cnf transformation 41] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
205. k6_numbers = k1_xboole_0 [cnf transformation 63] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
207. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 72] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
218. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 97] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
219. ~v7_ordinal1(X0) | v1_xcmplx_0(X0) [cnf transformation 98] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
221. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 100] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
222. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 101] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
223. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 102] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
224. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 104] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
225. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X1)) | r1_xxreal_0(k6_numbers,X1) | r1_xxreal_0(X0,k6_numbers) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 106] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
226. k6_numbers = k3_xcmplx_0(X0,k6_numbers) | ~v1_xcmplx_0(X0) [cnf transformation 107] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
234. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 114] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
235. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 115] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
269. v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 151] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
271. k3_xcmplx_0(X0,X1) = k3_xcmplx_0(X1,X0) | ~v1_xcmplx_0(X1) | ~v1_xcmplx_0(X0) [cnf transformation 155] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
273. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 159] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
297. 1 <=> r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) [avatar definition]
299. ~r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) <- (~1) [avatar component clause 297] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
300. ~1 [avatar split clause 200,297]
302. 2 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
304. r1_xxreal_0(k6_numbers,sK1) <- (2) [avatar component clause 302] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
305. 2 [avatar split clause 199,302]
307. 3 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
309. r1_xxreal_0(sK0,k6_numbers) <- (3) [avatar component clause 307] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
310. 3 [avatar split clause 198,307]
312. 4 <=> v1_xreal_0(sK1) [avatar definition]
314. v1_xreal_0(sK1) <- (4) [avatar component clause 312] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
315. 4 [avatar split clause 197,312]
317. 5 <=> v1_xreal_0(sK0) [avatar definition]
319. v1_xreal_0(sK0) <- (5) [avatar component clause 317] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
320. 5 [avatar split clause 196,317]
332. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
334. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 332] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
335. 8 [avatar split clause 203,332]
342. 10 <=> k6_numbers = k1_xboole_0 [avatar definition]
344. k6_numbers = k1_xboole_0 <- (10) [avatar component clause 342] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
345. 10 [avatar split clause 205,342]
352. 12 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
354. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (12) [avatar component clause 352] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
355. 12 [avatar split clause 207,352]
405. 23 <=> ! [X0] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0)) [avatar definition]
406. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) <- (23) [avatar component clause 405] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
513. v1_xboole_0(k6_numbers) <- (8, 10) [backward demodulation 334,344] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
515. 45 <=> v1_xboole_0(k6_numbers) [avatar definition]
517. v1_xboole_0(k6_numbers) <- (45) [avatar component clause 515] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
518. 45 | ~8 | ~10 [avatar split clause 513,342,332,515]
519. m1_subset_1(k6_numbers,k4_ordinal1) <- (10, 12) [forward demodulation 354,344] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
521. 46 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
523. m1_subset_1(k6_numbers,k4_ordinal1) <- (46) [avatar component clause 521] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
524. 46 | ~10 | ~12 [avatar split clause 519,352,342,521]
579. v1_xcmplx_0(sK0) <- (5) [resolution 222,319] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
585. 56 <=> v1_xcmplx_0(sK0) [avatar definition]
587. v1_xcmplx_0(sK0) <- (56) [avatar component clause 585] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
588. 56 | ~5 [avatar split clause 579,317,585]
590. 57 <=> v1_xcmplx_0(sK1) [avatar definition]
592. v1_xcmplx_0(sK1) <- (57) [avatar component clause 590] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
680. 71 <=> v7_ordinal1(k6_numbers) [avatar definition]
682. v7_ordinal1(k6_numbers) <- (71) [avatar component clause 680] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
744. k6_numbers = X0 | ~v1_xboole_0(X0) <- (10) [forward demodulation 234,344] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
745. v7_ordinal1(k6_numbers) <- (46) [resolution 235,523] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
748. 71 | ~46 [avatar split clause 745,521,680]
784. ~r1_xxreal_0(k3_xcmplx_0(sK0,sK1),X0) | ~v1_xboole_0(X0) <- (~1, 10) [superposition 299,744] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
817. 91 <=> v1_xxreal_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
818. v1_xxreal_0(k3_xcmplx_0(sK0,sK1)) <- (91) [avatar component clause 817] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
819. ~v1_xxreal_0(k3_xcmplx_0(sK0,sK1)) <- (~91) [avatar component clause 817] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
821. 92 <=> v1_xboole_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
823. ~v1_xboole_0(k3_xcmplx_0(sK0,sK1)) <- (~92) [avatar component clause 821] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
826. 93 <=> v1_xboole_0(sK1) [avatar definition]
831. k3_xcmplx_0(X1,X0) = X0 | ~v1_xcmplx_0(X1) | ~v1_xboole_0(X0) <- (10) [superposition 226,744] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1048. v1_xxreal_0(k3_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [resolution 269,223] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1070. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X0) [factoring 273] {a:1,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
1071. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) [duplicate literal removal 1070] {a:1,w:5,wCS:2000,sine:1,thAx:0,allAx:1,thDist:-1}
1073. 130 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1074. v1_xxreal_0(k6_numbers) <- (130) [avatar component clause 1073] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1094. k6_numbers = k3_xcmplx_0(k6_numbers,X2) | ~v1_xcmplx_0(X2) | ~v1_xcmplx_0(X2) | ~v1_xcmplx_0(k6_numbers) [superposition 226,271] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
1098. k6_numbers = k3_xcmplx_0(k6_numbers,X2) | ~v1_xcmplx_0(X2) | ~v1_xcmplx_0(k6_numbers) [duplicate literal removal 1094] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
1112. 134 <=> v1_xcmplx_0(k6_numbers) [avatar definition]
1116. 135 <=> ! [X1] : (k6_numbers = k3_xcmplx_0(k6_numbers,X1) | ~v1_xcmplx_0(X1)) [avatar definition]
1117. k6_numbers = k3_xcmplx_0(k6_numbers,X1) | ~v1_xcmplx_0(X1) <- (135) [avatar component clause 1116] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1124. ~134 | 135 [avatar split clause 1098,1116,1112]
1135. v1_xxreal_0(k6_numbers) <- (71) [resolution 682,221] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1137. v1_xcmplx_0(k6_numbers) <- (71) [resolution 682,219] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1138. v1_xreal_0(k6_numbers) <- (71) [resolution 682,218] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1140. 130 | ~71 [avatar split clause 1135,680,1073]
1142. 134 | ~71 [avatar split clause 1137,680,1112]
1144. 136 <=> v1_xreal_0(k6_numbers) [avatar definition]
1147. 136 | ~71 [avatar split clause 1138,680,1144]
1160. 137 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
1162. ~r1_xxreal_0(sK1,k6_numbers) <- (~137) [avatar component clause 1160] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1164. 138 <=> k6_numbers = sK1 [avatar definition]
1170. 139 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
1171. r1_xxreal_0(k6_numbers,sK0) <- (139) [avatar component clause 1170] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1174. 140 <=> k6_numbers = sK0 [avatar definition]
1176. k6_numbers = sK0 <- (140) [avatar component clause 1174] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1180. ~v1_xboole_0(X0) | v1_xxreal_0(X0) <- (10, 130) [superposition 1074,744] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1181. r1_xxreal_0(k6_numbers,X0) | r1_xxreal_0(X1,k6_numbers) | ~v1_xreal_0(X0) | ~v1_xreal_0(X1) | r1_xxreal_0(k3_xcmplx_0(X1,X0),k6_numbers) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(k3_xcmplx_0(X1,X0)) [resolution 225,273] {a:1,w:21,wCS:8400,sine:1,thAx:0,allAx:2,thDist:-2}
1186. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X3,X2)) | r1_xxreal_0(k6_numbers,X3) | r1_xxreal_0(X2,k6_numbers) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) | ~v1_xcmplx_0(X2) | ~v1_xcmplx_0(X3) [superposition 225,271] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:2,thDist:-2}
1189. r1_xxreal_0(k3_xcmplx_0(X1,X0),k6_numbers) | r1_xxreal_0(X1,k6_numbers) | ~v1_xreal_0(X0) | ~v1_xreal_0(X1) | r1_xxreal_0(k6_numbers,X0) | ~v1_xxreal_0(k3_xcmplx_0(X1,X0)) <- (130) [subsumption resolution 1181,1074] {a:1,w:19,wCS:1900,nSel:1,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1190. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X3,X2)) | r1_xxreal_0(k6_numbers,X3) | r1_xxreal_0(X2,k6_numbers) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) | ~v1_xcmplx_0(X3) [subsumption resolution 1186,222] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
1191. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X3,X2)) | r1_xxreal_0(k6_numbers,X3) | r1_xxreal_0(X2,k6_numbers) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) [subsumption resolution 1190,222] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1460. r1_xxreal_0(k6_numbers,k6_numbers) | r1_xxreal_0(X1,k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X1) | r1_xxreal_0(k6_numbers,k6_numbers) | ~v1_xxreal_0(k6_numbers) | ~v1_xcmplx_0(X1) <- (130) [superposition 1189,226] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1465. r1_xxreal_0(k6_numbers,k6_numbers) | r1_xxreal_0(X1,k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X1) | ~v1_xxreal_0(k6_numbers) | ~v1_xcmplx_0(X1) <- (130) [duplicate literal removal 1460] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1603. k6_numbers != sK1 | k6_numbers != k1_xboole_0 | v1_xboole_0(sK1) | ~v1_xboole_0(k1_xboole_0) [theory tautology sat conflict] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:1,thDist:-1}
1619. 175 <=> ! [X8] : (r1_xxreal_0(X8,k6_numbers) | ~v1_xreal_0(X8)) [avatar definition]
1620. r1_xxreal_0(X8,k6_numbers) | ~v1_xreal_0(X8) <- (175) [avatar component clause 1619] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1671. 178 <=> r1_xxreal_0(k6_numbers,k6_numbers) [avatar definition]
1673. r1_xxreal_0(k6_numbers,k6_numbers) <- (178) [avatar component clause 1671] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1687. v1_xcmplx_0(sK1) <- (4) [resolution 314,222] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1689. 57 | ~4 [avatar split clause 1687,312,590]
1690. k6_numbers = sK1 | ~r1_xxreal_0(sK1,k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (2) [resolution 304,224] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1692. k6_numbers = sK1 | ~r1_xxreal_0(sK1,k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 4) [subsumption resolution 1690,314] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1693. ~136 | ~137 | 138 | ~2 | ~4 [avatar split clause 1692,312,302,1164,1160,1144]
1718. ~v1_xboole_0(k3_xcmplx_0(sK0,sK1)) | ~v1_xxreal_0(k3_xcmplx_0(sK0,sK1)) <- (~1, 10, 23) [resolution 784,406] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1721. ~v1_xboole_0(k3_xcmplx_0(sK0,sK1)) <- (~1, 10, 23, 130) [subsumption resolution 1718,1180] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1722. ~92 | 1 | ~10 | ~23 | ~130 [avatar split clause 1721,1073,405,342,297,821]
1725. 181 <=> ! [X0] : (~v1_xboole_0(X0) | r1_xxreal_0(X0,k3_xcmplx_0(sK0,sK1))) [avatar definition]
1726. r1_xxreal_0(X0,k3_xcmplx_0(sK0,sK1)) | ~v1_xboole_0(X0) <- (181) [avatar component clause 1725] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1815. ~v1_xreal_0(sK0) | ~v1_xreal_0(sK1) <- (~91) [resolution 819,1048] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1817. ~v1_xreal_0(sK1) <- (5, ~91) [subsumption resolution 1815,319] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1818. $false <- (4, 5, ~91) [subsumption resolution 1817,314] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1819. ~4 | ~5 | 91 [avatar contradiction clause 1818]
1887. r1_xxreal_0(k6_numbers,k6_numbers) | r1_xxreal_0(X1,k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X1) | ~v1_xcmplx_0(X1) <- (130) [subsumption resolution 1465,1074] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1888. r1_xxreal_0(k6_numbers,k6_numbers) | r1_xxreal_0(X1,k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X1) <- (130) [subsumption resolution 1887,222] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1889. ~136 | 175 | 178 | ~130 [avatar split clause 1888,1073,1671,1619,1144]
1890. 23 [avatar split clause 1071,405]
1894. ~v1_xboole_0(sK1) | ~v1_xcmplx_0(sK0) | ~v1_xboole_0(sK1) <- (10, ~92) [superposition 823,831] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1895. ~v1_xboole_0(sK1) | ~v1_xcmplx_0(sK0) <- (10, ~92) [duplicate literal removal 1894] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1991. r1_xxreal_0(X5,k3_xcmplx_0(sK0,sK1)) | ~v1_xxreal_0(k3_xcmplx_0(sK0,sK1)) | ~v1_xxreal_0(X5) | ~v1_xboole_0(X5) <- (~1, 10) [resolution 273,784] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1998. r1_xxreal_0(X5,k3_xcmplx_0(sK0,sK1)) | ~v1_xxreal_0(X5) | ~v1_xboole_0(X5) <- (~1, 10, 91) [subsumption resolution 1991,818] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1999. r1_xxreal_0(X5,k3_xcmplx_0(sK0,sK1)) | ~v1_xboole_0(X5) <- (~1, 10, 91, 130) [subsumption resolution 1998,1180] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2000. 181 | 1 | ~10 | ~91 | ~130 [avatar split clause 1999,1073,817,342,297,1725]
2007. ~v1_xreal_0(sK1) <- (~137, 175) [resolution 1620,1162] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
2015. $false <- (4, ~137, 175) [subsumption resolution 2007,314] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
2016. ~4 | 137 | ~175 [avatar contradiction clause 2015]
2192. ~v1_xboole_0(sK1) <- (10, 56, ~92) [subsumption resolution 1895,587] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2193. ~93 | ~10 | ~56 | 92 [avatar split clause 2192,821,585,342,826]
2604. r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(sK1,k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(sK1) | ~v1_xboole_0(k6_numbers) <- (181) [resolution 1191,1726] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2645. r1_xxreal_0(k6_numbers,sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(sK1) | ~v1_xboole_0(k6_numbers) <- (~137, 181) [subsumption resolution 2604,1162] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2646. r1_xxreal_0(k6_numbers,sK0) | ~v1_xreal_0(sK1) | ~v1_xboole_0(k6_numbers) <- (5, ~137, 181) [subsumption resolution 2645,319] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2647. r1_xxreal_0(k6_numbers,sK0) | ~v1_xboole_0(k6_numbers) <- (4, 5, ~137, 181) [subsumption resolution 2646,314] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2648. r1_xxreal_0(k6_numbers,sK0) <- (4, 5, 45, ~137, 181) [subsumption resolution 2647,517] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2649. 139 | ~4 | ~5 | ~45 | 137 | ~181 [avatar split clause 2648,1725,1160,515,317,312,1170]
2674. k6_numbers = sK0 | ~r1_xxreal_0(sK0,k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (139) [resolution 1171,224] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2676. k6_numbers = sK0 | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (3, 139) [subsumption resolution 2674,309] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2677. k6_numbers = sK0 | ~v1_xreal_0(k6_numbers) <- (3, 5, 139) [subsumption resolution 2676,319] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2678. ~136 | 140 | ~3 | ~5 | ~139 [avatar split clause 2677,1170,317,307,1174,1144]
2679. ~r1_xxreal_0(k3_xcmplx_0(k6_numbers,sK1),k6_numbers) <- (~1, 140) [backward demodulation 299,1176] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
2698. 211 <=> r1_xxreal_0(k3_xcmplx_0(k6_numbers,sK1),k6_numbers) [avatar definition]
2700. ~r1_xxreal_0(k3_xcmplx_0(k6_numbers,sK1),k6_numbers) <- (~211) [avatar component clause 2698] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
2701. ~211 | 1 | ~140 [avatar split clause 2679,1174,297,2698]
2756. ~r1_xxreal_0(k6_numbers,k6_numbers) | ~v1_xcmplx_0(sK1) <- (135, ~211) [superposition 2700,1117] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2761. ~v1_xcmplx_0(sK1) <- (135, 178, ~211) [subsumption resolution 2756,1673] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2762. $false <- (57, 135, 178, ~211) [subsumption resolution 2761,592] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2763. ~57 | ~135 | ~178 | 211 [avatar contradiction clause 2762]
2764. $false [avatar sat refutation 300,305,310,315,320,335,345,355,518,524,588,748,1124,1140,1142,1147,1603,1689,1693,1722,1819,1889,1890,2000,2016,2193,2649,2678,2701,2763] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:2245,thDist:-2245}
% SZS output end Proof for xreal_0__l27_xreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 81
% >>> Preprocessing
% Introduced skolems: 17
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 100
% Generated clauses: 1672
% Active clauses: 267
% Passive clauses: 486
% Final active clauses: 189
% Final passive clauses: 181
% Main loop iterations started: 286
% 
% >>> Simplifying Inferences
% Duplicate literals: 192
% Fw subsumption resolutions: 269
% Fw demodulations: 15
% Bw demodulations: 32
% 
% >>> Deletion Inferences
% Simple tautologies: 72
% Equational tautologies: 5
% Forward subsumptions: 319
% Backward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 393
% Unit resulting resolution: 22
% Factoring: 2
% Forward superposition: 258
% Backward superposition: 118
% Self superposition: 5
% 
% >>> AVATAR
% Split clauses: 80
% Split components: 192
% Sat splitting refutations: 47
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 345
% SAT solver unit clauses: 46
% SAT solver binary clauses: 132
% 
% Memory used [KB]: 6652
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% term sharing: 0.002 s
% splitting model update: 0.001 s
% congruence closure: 0.002 s
% binary resolution index maintenance: 0.001 s
% backward subsumption index maintenance: 0.001 s
% resolution: 0.001 s
% unit resulting resolution: 0.002 s ( own 0.001 s ) 
% superposition: 0.004 s ( own 0.003 s ) 
% other: 0.016 s ( own 0.003 s ) 

