i: [1,36,0,11,3,1,0,0] t1_xboole_0
a: [3,66,0,3,1,36]
s: 66
a: [3,70,0,3,1,36]
s: 70
a: [3,63,0,5,1,36]
s: 63
i: [1,50,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 50
i: [1,53,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 53
i: [1,51,0,7,1,0,0,1] commutativity_k5_xboole_0
s: 51
d: [2,76,1,3,1,1,40,53,66]
a: [3,80,1,3,1,76]
s: 80
i: [1,52,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 52
i: [1,55,0,8,2,0,0,255] d5_xboole_0
s: 55
i: [1,56,0,8,2,0,0,255] d5_xboole_0
s: 56
i: [1,57,0,8,2,0,0,1] d3_xboole_0
s: 57
i: [1,58,0,8,2,0,0,1] d3_xboole_0
s: 58
i: [1,37,0,11,1,0,0,255] d6_xboole_0
s: 37
d: [2,90,1,7,1,0,44,37,50]
s: 90
i: [1,54,0,11,3,0,0,255] d5_xboole_0
s: 54
i: [1,59,0,11,3,0,0,1] d3_xboole_0
s: 59
d: [2,93,1,7,1,0,44,50,37]
s: 93
i: [1,41,0,17,3,0,0,255] d5_xboole_0
s: 41
d: [2,82,1,8,2,0,40,57,53]
s: 82
i: [1,42,0,17,3,0,0,255] d5_xboole_0
s: 42
i: [1,48,0,17,3,0,0,1] d3_xboole_0
s: 48
d: [2,84,1,8,2,0,44,57,52]
s: 84
i: [1,49,0,17,3,0,0,1] d3_xboole_0
s: 49
d: [2,85,1,8,2,0,44,57,52]
s: 85
i: [1,43,0,23,4,0,0,255] d5_xboole_0
s: 43
i: [1,47,0,23,4,0,0,1] d3_xboole_0
s: 47
d: [2,86,1,8,2,0,40,58,53]
s: 86
d: [2,94,1,10,2,0,44,58,37]
s: 94
d: [2,100,2,8,2,0,44,56,90]
s: 100
d: [2,95,1,10,2,0,44,57,37]
s: 95
d: [2,91,1,11,1,0,44,37,52]
s: 91
d: [2,101,2,8,2,0,44,55,90]
s: 101
d: [2,92,1,11,1,0,44,37,52]
s: 92
d: [2,185,2,5,1,0,21,101,100]
s: 185
r: 100
d: [2,96,1,11,1,0,44,52,37]
s: 96
d: [2,97,1,11,1,0,44,52,37]
s: 97
d: [2,122,2,8,2,0,44,82,52]
s: 122
d: [2,104,1,11,3,0,40,54,53]
s: 104
d: [2,123,2,8,2,0,44,82,52]
s: 123
d: [2,117,1,11,2,0,42,41]
s: 117
d: [2,147,1,17,3,0,42,47]
d: [2,151,1,11,2,0,21,147,49]
s: 151
d: [2,207,3,7,1,0,40,151,185]
s: 207
d: [2,148,1,17,3,0,42,47]
d: [2,152,1,11,2,0,21,148,48]
s: 152
d: [2,227,3,7,1,0,40,152,185]
s: 227
d: [2,111,1,15,3,0,44,59,37]
s: 111
d: [2,113,1,17,3,0,40,41,53]
s: 113
d: [2,215,4,7,1,0,44,207,52]
s: 215
d: [2,114,1,23,3,0,40,41,56]
s: 114
d: [2,216,4,7,1,0,44,207,52]
s: 216
d: [2,115,1,23,3,0,40,41,55]
s: 115
d: [2,126,1,23,3,0,40,42,58]
s: 126
d: [2,217,4,7,1,0,44,52,207]
s: 217
d: [2,127,1,23,3,0,40,42,57]
s: 127
d: [2,218,4,7,1,0,44,52,207]
s: 218
d: [2,130,1,23,3,0,40,48,58]
s: 130
d: [2,131,1,23,3,0,40,48,57]
s: 131
d: [2,237,4,7,1,0,44,227,207]
s: 237
d: [2,133,1,23,3,0,40,49,58]
s: 133
d: [2,245,4,7,1,0,44,207,227]
s: 245
d: [2,134,1,23,3,0,40,49,57]
s: 134
d: [2,143,1,23,4,0,40,47,53]
s: 143
d: [2,305,2,28,3,0,40,130,151]
d: [2,318,2,19,2,0,18,305]
d: [2,319,2,9,1,0,21,318,151]
s: 319
d: [2,116,1,31,4,0,40,41,59]
s: 116
d: [2,321,2,28,3,0,40,131,151]
d: [2,334,2,19,2,0,18,321]
d: [2,335,2,9,1,0,21,334,151]
s: 335
d: [2,128,1,31,4,0,40,42,54]
s: 128
d: [2,132,1,31,4,0,40,48,54]
s: 132
d: [2,352,2,28,3,0,40,133,152]
d: [2,363,2,19,2,0,18,352]
d: [2,365,2,9,1,0,21,363,152]
s: 365
d: [2,135,1,31,4,0,40,49,54]
s: 135
d: [2,367,2,28,3,0,40,134,152]
d: [2,378,2,19,2,0,18,367]
d: [2,380,2,9,1,0,21,378,152]
s: 380
d: [2,137,1,31,4,0,40,43,58]
s: 137
d: [2,138,1,31,4,0,40,43,57]
s: 138
d: [2,195,3,9,1,0,40,117,185]
s: 195
d: [2,144,1,31,4,0,40,47,56]
s: 144
d: [2,382,3,9,1,0,44,319,52]
s: 382
d: [2,145,1,31,4,0,40,47,55]
s: 145
d: [2,139,1,39,5,0,40,43,54]
s: 139
d: [2,383,3,9,1,0,44,319,52]
s: 383
d: [2,146,1,39,5,0,40,47,59]
s: 146
d: [2,477,3,9,1,0,44,365,52]
s: 477
d: [2,118,2,10,2,0,40,82,58]
s: 118
d: [2,119,2,10,2,0,40,82,57]
s: 119
d: [2,478,3,9,1,0,44,365,52]
s: 478
d: [2,124,2,10,2,0,44,82,37]
s: 124
d: [2,481,3,9,1,0,44,365,335]
s: 481
d: [2,153,2,10,2,0,40,86,58]
s: 153
d: [2,154,2,10,2,0,40,86,57]
s: 154
d: [2,159,2,10,2,0,44,86,37]
s: 159
d: [2,183,2,10,2,0,44,82,91]
s: 183
d: [2,184,2,10,2,0,44,86,91]
s: 184
d: [2,160,2,11,3,0,40,94,54]
s: 160
d: [2,166,2,11,3,0,40,95,54]
s: 166
d: [2,771,3,0,0,3,52,66,70,63,166]
e: 771
d: [2,783,3,6,2,1,40,166,63]
d: [2,807,3,3,1,2,21,783,70]
d: [2,808,3,0,0,3,21,807,66]
e: 808
r: 63
r: 66
a: [3,67,0,3,1,36]
s: 67
a: [3,62,0,5,1,36]
s: 62
d: [2,811,1,5,1,1,40,62,53]
a: [3,824,1,5,1,811]
s: 824
d: [2,810,2,10,2,1,40,62,111]
a: [3,819,2,5,1,810]
s: 819
d: [2,826,3,0,0,2,52,70,819,56]
e: 826
d: [2,829,3,3,1,1,40,819,56]
d: [2,834,3,0,0,2,21,829,70]
e: 834
r: 819
a: [3,815,2,5,1,810]
s: 815
d: [2,836,3,0,0,2,52,67,815,56]
e: 836
d: [2,839,3,3,1,1,40,815,56]
d: [2,844,3,0,0,2,21,839,67]
e: 844
r: 67
r: 70
r: 815
s: 66
a: [3,71,0,3,1,36]
s: 71
d: [2,847,1,3,1,1,40,71,53]
a: [3,851,1,3,1,847]
s: 851
s: 819
d: [2,853,3,0,0,2,52,66,819,55]
e: 853
d: [2,857,3,3,1,1,40,819,55]
d: [2,861,3,0,0,2,21,857,66]
e: 861
r: 819
s: 815
d: [2,863,3,0,0,2,52,71,815,55]
e: 863
d: [2,867,3,3,1,1,40,815,55]
d: [2,871,3,0,0,2,21,867,71]
e: 871
r: 62
r: 66
r: 815
s: 67
s: 63
d: [2,873,3,0,0,3,52,71,67,63,160]
e: 873
d: [2,875,3,6,2,1,40,63,160]
d: [2,877,3,3,1,2,21,875,67]
d: [2,878,3,0,0,3,21,877,71]
e: 878
i: [1,35,0,11,3,1,0,0] t1_xboole_0
i: [1,34,0,11,3,1,0,0] t1_xboole_0
i: [1,33,0,11,3,1,0,0] t1_xboole_0
f: 878,871,861,844,834,810,808,33,34,35,36
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_0__t1_xboole_0
% SZS output start Proof for xboole_0__t1_xboole_0
1. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k5_xboole_0(X1,X2)) <=> ~(r2_hidden(X0,X1) <=> r2_hidden(X0,X2))) [input t1_xboole_0]
2. ~! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k5_xboole_0(X1,X2)) <=> ~(r2_hidden(X0,X1) <=> r2_hidden(X0,X2))) [negated conjecture 1]
6. ! [X0] : ! [X1] : ! [X2] : (k2_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) | r2_hidden(X3,X0)))) [input d3_xboole_0]
7. ! [X0] : ! [X1] : ! [X2] : (k4_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (~r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [input d5_xboole_0]
8. ! [X0] : ! [X1] : k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [input d6_xboole_0]
13. ~! [X0,X1,X2] : (r2_hidden(X0,k5_xboole_0(X1,X2)) <=> ~(r2_hidden(X0,X1) <=> r2_hidden(X0,X2))) [flattening 2]
14. ! [X0,X1] : k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [flattening 8]
15. ! [X0,X1,X2] : (k4_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (~r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [flattening 7]
16. ! [X0,X1,X2] : (k2_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) | r2_hidden(X3,X0)))) [flattening 6]
18. ? [X0,X1,X2] : (r2_hidden(X0,k5_xboole_0(X1,X2)) <~> (r2_hidden(X0,X1) <~> r2_hidden(X0,X2))) [ennf transformation 13]
20. ? [X0,X1,X2] : ((((r2_hidden(X0,X1) | ~r2_hidden(X0,X2)) & (r2_hidden(X0,X2) | ~r2_hidden(X0,X1))) | ~r2_hidden(X0,k5_xboole_0(X1,X2))) & (((~r2_hidden(X0,X2) | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X2) | r2_hidden(X0,X1))) | r2_hidden(X0,k5_xboole_0(X1,X2)))) [nnf transformation 18]
21. ? [X0,X1,X2] : ((((r2_hidden(X0,X1) | ~r2_hidden(X0,X2)) & (r2_hidden(X0,X2) | ~r2_hidden(X0,X1))) | ~r2_hidden(X0,k5_xboole_0(X1,X2))) & (((~r2_hidden(X0,X2) | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X2) | r2_hidden(X0,X1))) | r2_hidden(X0,k5_xboole_0(X1,X2)))) => ((((r2_hidden(sK0,sK1) | ~r2_hidden(sK0,sK2)) & (r2_hidden(sK0,sK2) | ~r2_hidden(sK0,sK1))) | ~r2_hidden(sK0,k5_xboole_0(sK1,sK2))) & (((~r2_hidden(sK0,sK2) | ~r2_hidden(sK0,sK1)) & (r2_hidden(sK0,sK2) | r2_hidden(sK0,sK1))) | r2_hidden(sK0,k5_xboole_0(sK1,sK2)))) [choice axiom]
22. (((r2_hidden(sK0,sK1) | ~r2_hidden(sK0,sK2)) & (r2_hidden(sK0,sK2) | ~r2_hidden(sK0,sK1))) | ~r2_hidden(sK0,k5_xboole_0(sK1,sK2))) & (((~r2_hidden(sK0,sK2) | ~r2_hidden(sK0,sK1)) & (r2_hidden(sK0,sK2) | r2_hidden(sK0,sK1))) | r2_hidden(sK0,k5_xboole_0(sK1,sK2))) [skolemisation 20,21]
23. ! [X0,X1,X2] : ((k4_xboole_0(X0,X1) = X2 | ? [X3] : (((r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (r2_hidden(X3,X1) | ~r2_hidden(X3,X0))) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k4_xboole_0(X0,X1) != X2)) [nnf transformation 15]
24. ! [X0,X1,X2] : ((k4_xboole_0(X0,X1) = X2 | ? [X3] : ((r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k4_xboole_0(X0,X1) != X2)) [flattening 23]
25. ! [X0,X1,X2] : ((k4_xboole_0(X0,X1) = X2 | ? [X3] : ((r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((~r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k4_xboole_0(X0,X1) != X2)) [rectify 24]
26. ! [X2,X1,X0] : (? [X3] : ((r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2))) => ((r2_hidden(sK3(X0,X1,X2),X1) | ~r2_hidden(sK3(X0,X1,X2),X0) | ~r2_hidden(sK3(X0,X1,X2),X2)) & ((~r2_hidden(sK3(X0,X1,X2),X1) & r2_hidden(sK3(X0,X1,X2),X0)) | r2_hidden(sK3(X0,X1,X2),X2)))) [choice axiom]
27. ! [X0,X1,X2] : ((k4_xboole_0(X0,X1) = X2 | ((r2_hidden(sK3(X0,X1,X2),X1) | ~r2_hidden(sK3(X0,X1,X2),X0) | ~r2_hidden(sK3(X0,X1,X2),X2)) & ((~r2_hidden(sK3(X0,X1,X2),X1) & r2_hidden(sK3(X0,X1,X2),X0)) | r2_hidden(sK3(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((~r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k4_xboole_0(X0,X1) != X2)) [skolemisation 25,26]
28. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) | r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) & ~r2_hidden(X3,X0))) & ((r2_hidden(X3,X1) | r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k2_xboole_0(X0,X1) != X2)) [nnf transformation 16]
29. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) & ~r2_hidden(X3,X0))) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | ~r2_hidden(X3,X2))) | k2_xboole_0(X0,X1) != X2)) [flattening 28]
30. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (~r2_hidden(X4,X1) & ~r2_hidden(X4,X0))) & (r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2))) | k2_xboole_0(X0,X1) != X2)) [rectify 29]
31. ! [X2,X1,X0] : (? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2))) => (((~r2_hidden(sK4(X0,X1,X2),X1) & ~r2_hidden(sK4(X0,X1,X2),X0)) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (r2_hidden(sK4(X0,X1,X2),X1) | r2_hidden(sK4(X0,X1,X2),X0) | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
32. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | (((~r2_hidden(sK4(X0,X1,X2),X1) & ~r2_hidden(sK4(X0,X1,X2),X0)) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (r2_hidden(sK4(X0,X1,X2),X1) | r2_hidden(sK4(X0,X1,X2),X0) | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (~r2_hidden(X4,X1) & ~r2_hidden(X4,X0))) & (r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2))) | k2_xboole_0(X0,X1) != X2)) [skolemisation 30,31]
33. r2_hidden(sK0,sK2) | r2_hidden(sK0,sK1) | r2_hidden(sK0,k5_xboole_0(sK1,sK2)) [cnf transformation 22] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. ~r2_hidden(sK0,sK2) | ~r2_hidden(sK0,sK1) | r2_hidden(sK0,k5_xboole_0(sK1,sK2)) [cnf transformation 22] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. r2_hidden(sK0,sK2) | ~r2_hidden(sK0,sK1) | ~r2_hidden(sK0,k5_xboole_0(sK1,sK2)) [cnf transformation 22] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. r2_hidden(sK0,sK1) | ~r2_hidden(sK0,sK2) | ~r2_hidden(sK0,k5_xboole_0(sK1,sK2)) [cnf transformation 22] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. k5_xboole_0(X0,X1) = k2_xboole_0(k4_xboole_0(X0,X1),k4_xboole_0(X1,X0)) [cnf transformation 14] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. r2_hidden(X4,X0) | ~r2_hidden(X4,X2) | k4_xboole_0(X0,X1) != X2 [cnf transformation 27] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
39. ~r2_hidden(X4,X1) | ~r2_hidden(X4,X2) | k4_xboole_0(X0,X1) != X2 [cnf transformation 27] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
40. r2_hidden(X4,X2) | r2_hidden(X4,X1) | ~r2_hidden(X4,X0) | k4_xboole_0(X0,X1) != X2 [cnf transformation 27] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:2,thDist:-2}
44. r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2) | k2_xboole_0(X0,X1) != X2 [cnf transformation 32] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
45. r2_hidden(X4,X2) | ~r2_hidden(X4,X0) | k2_xboole_0(X0,X1) != X2 [cnf transformation 32] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
46. r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | k2_xboole_0(X0,X1) != X2 [cnf transformation 32] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
54. r2_hidden(X4,k4_xboole_0(X0,X1)) | r2_hidden(X4,X1) | ~r2_hidden(X4,X0) [equality resolution 40] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
55. ~r2_hidden(X4,k4_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) [equality resolution 39] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. ~r2_hidden(X4,k4_xboole_0(X0,X1)) | r2_hidden(X4,X0) [equality resolution 38] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
57. r2_hidden(X4,k2_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) [equality resolution 46] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
58. r2_hidden(X4,k2_xboole_0(X0,X1)) | ~r2_hidden(X4,X0) [equality resolution 45] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
59. ~r2_hidden(X4,k2_xboole_0(X0,X1)) | r2_hidden(X4,X0) | r2_hidden(X4,X1) [equality resolution 44] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
61. 1 <=> r2_hidden(sK0,k5_xboole_0(sK1,sK2)) [avatar definition]
62. r2_hidden(sK0,k5_xboole_0(sK1,sK2)) <- (1) [avatar component clause 61] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. ~r2_hidden(sK0,k5_xboole_0(sK1,sK2)) <- (~1) [avatar component clause 61] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. 2 <=> r2_hidden(sK0,sK2) [avatar definition]
66. r2_hidden(sK0,sK2) <- (2) [avatar component clause 65] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
67. ~r2_hidden(sK0,sK2) <- (~2) [avatar component clause 65] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
69. 3 <=> r2_hidden(sK0,sK1) [avatar definition]
70. ~r2_hidden(sK0,sK1) <- (~3) [avatar component clause 69] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. r2_hidden(sK0,sK1) <- (3) [avatar component clause 69] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. ~1 | ~2 | 3 [avatar split clause 36,69,65,61]
73. ~1 | ~3 | 2 [avatar split clause 35,65,69,61]
74. 1 | ~3 | ~2 [avatar split clause 34,65,69,61]
75. 1 | 3 | 2 [avatar split clause 33,65,69,61]
94. ~r2_hidden(X3,k4_xboole_0(X1,X2)) | r2_hidden(X3,k5_xboole_0(X1,X2)) [superposition 58,37] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
95. ~r2_hidden(X6,k4_xboole_0(X5,X4)) | r2_hidden(X6,k5_xboole_0(X4,X5)) [superposition 57,37] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
111. ~r2_hidden(X10,k5_xboole_0(X8,X9)) | r2_hidden(X10,k4_xboole_0(X8,X9)) | r2_hidden(X10,k4_xboole_0(X9,X8)) [superposition 59,37] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
160. r2_hidden(X0,k5_xboole_0(X1,X2)) | r2_hidden(X0,X2) | ~r2_hidden(X0,X1) [resolution 94,54] {a:2,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
166. r2_hidden(X0,k5_xboole_0(X1,X2)) | r2_hidden(X0,X1) | ~r2_hidden(X0,X2) [resolution 95,54] {a:2,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
783. r2_hidden(sK0,sK1) | ~r2_hidden(sK0,sK2) <- (~1) [resolution 166,63] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
807. ~r2_hidden(sK0,sK2) <- (~1, ~3) [subsumption resolution 783,70] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
808. $false <- (~1, 2, ~3) [subsumption resolution 807,66] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
809. 1 | ~2 | 3 [avatar contradiction clause 808]
810. r2_hidden(sK0,k4_xboole_0(sK1,sK2)) | r2_hidden(sK0,k4_xboole_0(sK2,sK1)) <- (1) [resolution 62,111] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
813. 5 <=> r2_hidden(sK0,k4_xboole_0(sK2,sK1)) [avatar definition]
815. r2_hidden(sK0,k4_xboole_0(sK2,sK1)) <- (5) [avatar component clause 813] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
817. 6 <=> r2_hidden(sK0,k4_xboole_0(sK1,sK2)) [avatar definition]
819. r2_hidden(sK0,k4_xboole_0(sK1,sK2)) <- (6) [avatar component clause 817] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
820. 5 | 6 | ~1 [avatar split clause 810,61,817,813]
829. r2_hidden(sK0,sK1) <- (6) [resolution 819,56] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
834. $false <- (~3, 6) [subsumption resolution 829,70] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
835. 3 | ~6 [avatar contradiction clause 834]
839. r2_hidden(sK0,sK2) <- (5) [resolution 815,56] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
844. $false <- (~2, 5) [subsumption resolution 839,67] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
845. 2 | ~5 [avatar contradiction clause 844]
857. ~r2_hidden(sK0,sK2) <- (6) [resolution 819,55] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
861. $false <- (2, 6) [subsumption resolution 857,66] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
862. ~2 | ~6 [avatar contradiction clause 861]
867. ~r2_hidden(sK0,sK1) <- (5) [resolution 815,55] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
871. $false <- (3, 5) [subsumption resolution 867,71] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
872. ~3 | ~5 [avatar contradiction clause 871]
875. r2_hidden(sK0,sK2) | ~r2_hidden(sK0,sK1) <- (~1) [resolution 63,160] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
877. ~r2_hidden(sK0,sK1) <- (~1, ~2) [subsumption resolution 875,67] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
878. $false <- (~1, ~2, 3) [subsumption resolution 877,71] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
879. 1 | 2 | ~3 [avatar contradiction clause 878]
880. $false [avatar sat refutation 72,73,74,75,809,820,835,845,862,872,879] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:143,thDist:-143}
% SZS output end Proof for xboole_0__t1_xboole_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 832
% Active clauses: 81
% Passive clauses: 358
% Final active clauses: 70
% Final passive clauses: 255
% Main loop iterations started: 103
% 
% >>> Simplifying Inferences
% Duplicate literals: 57
% Fw subsumption resolutions: 22
% Fw demodulations: 34
% 
% >>> Deletion Inferences
% Simple tautologies: 47
% Forward subsumptions: 288
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 297
% Unit resulting resolution: 6
% Factoring: 11
% Forward superposition: 222
% Backward superposition: 114
% Self superposition: 8
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 14
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 22
% SAT solver binary clauses: 11
% 
% Memory used [KB]: 9978
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.002 s
% superposition: 0.005 s
% other: 0.011 s ( own 0.004 s ) 

