i: [1,54,0,2,1,1,0,0] t43_funct_1
a: [3,101,0,2,1,54]
s: 101
i: [1,53,0,2,1,1,0,0] t43_funct_1
a: [3,106,0,2,1,53]
s: 106
i: [1,52,0,2,1,1,0,0] t43_funct_1
a: [3,111,0,2,1,52]
s: 111
i: [1,55,0,5,1,1,0,0] t43_funct_1
a: [3,96,0,5,1,55]
s: 96
i: [1,89,0,2,1,0,0,1] rc1_funct_1
a: [3,116,0,2,1,89]
s: 116
i: [1,88,0,2,1,0,0,1] rc1_funct_1
a: [3,121,0,2,1,88]
s: 121
i: [1,92,0,2,1,0,0,1] rc2_funct_1
a: [3,126,0,2,1,92]
s: 126
i: [1,91,0,2,1,0,0,1] rc2_funct_1
a: [3,131,0,2,1,91]
s: 131
i: [1,90,0,2,1,0,0,1] rc2_funct_1
a: [3,136,0,2,1,90]
s: 136
i: [1,56,0,3,1,0,0,2] dt_k4_relat_1
s: 56
i: [1,59,0,3,1,0,0,1] fc4_funct_1
s: 59
i: [1,60,0,3,1,0,0,1] fc4_funct_1
s: 60
i: [1,61,0,3,1,0,0,2] fc3_funct_1
s: 61
i: [1,62,0,3,1,0,0,2] fc3_funct_1
s: 62
i: [1,63,0,3,1,0,0,255] fc28_relat_1
s: 63
i: [1,64,0,4,1,0,0,255] fc28_relat_1
s: 64
i: [1,65,0,4,1,0,0,255] fc28_relat_1
s: 65
i: [1,76,0,4,1,0,0,255] dt_k3_relat_1
s: 76
i: [1,81,0,4,1,0,0,255] rc3_relat_1
s: 81
i: [1,57,0,5,1,0,0,255] rd2_relat_1
s: 57
i: [1,58,0,5,1,0,0,255] rd1_relat_1
s: 58
i: [1,82,0,5,1,0,0,255] rc3_relat_1
s: 82
i: [1,83,0,5,1,0,0,255] rc3_relat_1
s: 83
i: [1,69,0,7,3,0,0,1] dt_k2_funct_1
s: 69
i: [1,70,0,7,3,0,0,1] dt_k2_funct_1
s: 70
i: [1,66,0,9,4,0,0,1] fc6_funct_1
s: 66
i: [1,67,0,9,4,0,0,1] fc6_funct_1
s: 67
i: [1,68,0,9,4,0,0,1] fc6_funct_1
s: 68
i: [1,84,0,11,4,0,0,255] fc30_relat_1
s: 84
i: [1,86,0,11,4,0,0,255] fc29_relat_1
s: 86
i: [1,71,0,12,4,0,0,255] t33_funct_1
s: 71
i: [1,72,0,12,4,0,0,255] t33_funct_1
s: 72
i: [1,79,0,12,5,0,0,255] fc2_funct_1
s: 79
i: [1,80,0,12,5,0,0,255] fc2_funct_1
s: 80
i: [1,85,0,12,4,0,0,255] fc30_relat_1
s: 85
i: [1,87,0,12,4,0,0,255] fc29_relat_1
s: 87
i: [1,73,0,14,4,0,0,255] t39_funct_1
s: 73
i: [1,74,0,14,4,0,0,255] t39_funct_1
s: 74
d: [2,139,1,19,7,0,44,85,73]
d: [2,143,1,17,6,0,18,139]
d: [2,146,1,14,5,0,21,143,69]
s: 146
i: [1,77,0,16,7,0,0,255] fc7_funct_1
s: 77
i: [1,78,0,16,7,0,0,255] fc7_funct_1
s: 78
d: [2,147,1,19,7,0,44,87,74]
d: [2,153,1,17,6,0,18,147]
d: [2,154,1,14,5,0,21,153,69]
s: 154
i: [1,75,0,26,8,0,0,255] t41_funct_1
s: 75
d: [2,138,1,20,7,0,44,87,73]
d: [2,144,1,18,6,0,18,138]
d: [2,145,1,15,5,0,21,144,69]
s: 145
d: [2,148,1,20,7,0,44,85,74]
d: [2,152,1,18,6,0,18,148]
d: [2,155,1,15,5,0,21,152,69]
s: 155
d: [2,171,1,38,11,0,44,75,74]
d: [2,172,1,34,9,0,18,171]
d: [2,173,1,28,8,0,21,172,72]
d: [2,174,1,25,7,0,21,173,68]
d: [2,175,1,22,6,0,21,174,70]
d: [2,176,1,19,5,0,21,175,69]
s: 176
d: [2,177,2,18,5,0,44,145,58]
d: [2,179,2,15,4,0,21,177,56]
d: [2,180,2,12,3,0,21,179,60]
d: [2,181,2,9,2,0,21,180,62]
s: 181
d: [2,185,2,18,5,0,44,155,57]
d: [2,187,2,15,4,0,21,185,56]
d: [2,188,2,12,3,0,21,187,60]
d: [2,189,2,9,2,0,21,188,62]
s: 189
d: [2,193,2,24,8,0,44,176,71]
d: [2,194,2,18,5,0,18,193]
d: [2,195,2,11,4,0,19,194]
s: 195
r: 176
d: [2,196,3,0,0,4,52,101,111,106,96,195]
e: 196
d: [2,197,3,9,4,1,44,96,195]
d: [2,205,3,6,3,1,19,197]
d: [2,207,3,4,2,2,21,205,106]
d: [2,208,3,2,1,3,21,207,111]
d: [2,209,3,0,0,4,21,208,101]
e: 209
f: 209,52,53,54,55
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_1__t43_funct_1
% SZS output start Proof for funct_1__t43_funct_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => k2_funct_1(k2_funct_1(X0)) = X0)) [input t43_funct_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => k2_funct_1(k2_funct_1(X0)) = X0)) [negated conjecture 1]
4. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0)))) [input dt_k2_funct_1]
14. ! [X0] : ((v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(k2_funct_1(X0)) & v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0)))) [input fc6_funct_1]
21. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))))) [input t33_funct_1]
22. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k3_relat_1(k2_funct_1(X0),X0) = k4_relat_1(k10_xtuple_0(X0)) & k3_relat_1(X0,k2_funct_1(X0)) = k4_relat_1(k9_xtuple_0(X0))))) [input t39_funct_1]
23. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ((k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(X1) & v2_funct_1(X0)) => k2_funct_1(X0) = X1))) [input t41_funct_1]
24. ? [X0] : ((k2_funct_1(k2_funct_1(X0)) != X0 & v2_funct_1(X0)) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
25. ? [X0] : (k2_funct_1(k2_funct_1(X0)) != X0 & v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 24]
26. ! [X0] : ((v2_funct_1(k2_funct_1(X0)) & v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | (~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 14]
27. ! [X0] : ((v2_funct_1(k2_funct_1(X0)) & v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 26]
28. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 4]
29. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 28]
30. ! [X0] : (((k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 21]
31. ! [X0] : ((k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 30]
32. ! [X0] : (((k3_relat_1(k2_funct_1(X0),X0) = k4_relat_1(k10_xtuple_0(X0)) & k3_relat_1(X0,k2_funct_1(X0)) = k4_relat_1(k9_xtuple_0(X0))) | ~v2_funct_1(X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 22]
33. ! [X0] : ((k3_relat_1(k2_funct_1(X0),X0) = k4_relat_1(k10_xtuple_0(X0)) & k3_relat_1(X0,k2_funct_1(X0)) = k4_relat_1(k9_xtuple_0(X0))) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 32]
34. ! [X0] : (! [X1] : ((k2_funct_1(X0) = X1 | (k3_relat_1(X0,X1) != k4_relat_1(k9_xtuple_0(X0)) | k10_xtuple_0(X0) != k9_xtuple_0(X1) | ~v2_funct_1(X0))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 23]
35. ! [X0] : (! [X1] : (k2_funct_1(X0) = X1 | k3_relat_1(X0,X1) != k4_relat_1(k9_xtuple_0(X0)) | k10_xtuple_0(X0) != k9_xtuple_0(X1) | ~v2_funct_1(X0) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 34]
44. ? [X0] : (k2_funct_1(k2_funct_1(X0)) != X0 & v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (sK0 != k2_funct_1(k2_funct_1(sK0)) & v2_funct_1(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
45. sK0 != k2_funct_1(k2_funct_1(sK0)) & v2_funct_1(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 25,44]
52. v1_relat_1(sK0) [cnf transformation 45] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. v1_funct_1(sK0) [cnf transformation 45] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. v2_funct_1(sK0) [cnf transformation 45] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. sK0 != k2_funct_1(k2_funct_1(sK0)) [cnf transformation 45] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
68. v2_funct_1(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. v1_relat_1(k2_funct_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. v1_funct_1(k2_funct_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
72. k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
74. k3_relat_1(k2_funct_1(X0),X0) = k4_relat_1(k10_xtuple_0(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
75. k3_relat_1(X0,X1) != k4_relat_1(k9_xtuple_0(X0)) | k2_funct_1(X0) = X1 | k10_xtuple_0(X0) != k9_xtuple_0(X1) | ~v2_funct_1(X0) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:26,wCS:10400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
94. 1 <=> sK0 = k2_funct_1(k2_funct_1(sK0)) [avatar definition]
96. sK0 != k2_funct_1(k2_funct_1(sK0)) <- (~1) [avatar component clause 94] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
97. ~1 [avatar split clause 55,94]
99. 2 <=> v2_funct_1(sK0) [avatar definition]
101. v2_funct_1(sK0) <- (2) [avatar component clause 99] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
102. 2 [avatar split clause 54,99]
104. 3 <=> v1_funct_1(sK0) [avatar definition]
106. v1_funct_1(sK0) <- (3) [avatar component clause 104] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
107. 3 [avatar split clause 53,104]
109. 4 <=> v1_relat_1(sK0) [avatar definition]
111. v1_relat_1(sK0) <- (4) [avatar component clause 109] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
112. 4 [avatar split clause 52,109]
171. k4_relat_1(k10_xtuple_0(X1)) != k4_relat_1(k9_xtuple_0(k2_funct_1(X1))) | k2_funct_1(k2_funct_1(X1)) = X1 | k9_xtuple_0(X1) != k10_xtuple_0(k2_funct_1(X1)) | ~v2_funct_1(k2_funct_1(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(k2_funct_1(X1)) | ~v1_relat_1(k2_funct_1(X1)) | ~v2_funct_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [superposition 75,74] {a:1,w:38,wCS:15200,sine:255,thAx:0,allAx:2,thDist:-2}
172. k4_relat_1(k10_xtuple_0(X1)) != k4_relat_1(k9_xtuple_0(k2_funct_1(X1))) | k2_funct_1(k2_funct_1(X1)) = X1 | k9_xtuple_0(X1) != k10_xtuple_0(k2_funct_1(X1)) | ~v2_funct_1(k2_funct_1(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(k2_funct_1(X1)) | ~v1_relat_1(k2_funct_1(X1)) | ~v2_funct_1(X1) [duplicate literal removal 171] {a:1,w:34,wCS:13600,sine:255,thAx:0,allAx:2,thDist:-2}
173. k4_relat_1(k10_xtuple_0(X1)) != k4_relat_1(k9_xtuple_0(k2_funct_1(X1))) | k2_funct_1(k2_funct_1(X1)) = X1 | ~v2_funct_1(k2_funct_1(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(k2_funct_1(X1)) | ~v1_relat_1(k2_funct_1(X1)) | ~v2_funct_1(X1) [subsumption resolution 172,72] {a:1,w:28,wCS:11200,sine:255,thAx:0,allAx:2,thDist:-2}
174. k4_relat_1(k10_xtuple_0(X1)) != k4_relat_1(k9_xtuple_0(k2_funct_1(X1))) | k2_funct_1(k2_funct_1(X1)) = X1 | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(k2_funct_1(X1)) | ~v1_relat_1(k2_funct_1(X1)) | ~v2_funct_1(X1) [subsumption resolution 173,68] {a:1,w:25,wCS:10000,sine:1,thAx:0,allAx:2,thDist:-2}
175. k4_relat_1(k10_xtuple_0(X1)) != k4_relat_1(k9_xtuple_0(k2_funct_1(X1))) | k2_funct_1(k2_funct_1(X1)) = X1 | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_relat_1(k2_funct_1(X1)) | ~v2_funct_1(X1) [subsumption resolution 174,70] {a:1,w:22,wCS:8800,sine:1,thAx:0,allAx:2,thDist:-2}
176. k4_relat_1(k10_xtuple_0(X1)) != k4_relat_1(k9_xtuple_0(k2_funct_1(X1))) | k2_funct_1(k2_funct_1(X1)) = X1 | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v2_funct_1(X1) [subsumption resolution 175,69] {a:1,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
193. k4_relat_1(k10_xtuple_0(X0)) != k4_relat_1(k10_xtuple_0(X0)) | k2_funct_1(k2_funct_1(X0)) = X0 | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v2_funct_1(X0) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [superposition 176,71] {a:2,w:24,wCS:9600,sine:1,thAx:0,allAx:3,thDist:-3}
194. k4_relat_1(k10_xtuple_0(X0)) != k4_relat_1(k10_xtuple_0(X0)) | k2_funct_1(k2_funct_1(X0)) = X0 | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v2_funct_1(X0) [duplicate literal removal 193] {a:2,w:18,wCS:7200,sine:1,thAx:0,allAx:3,thDist:-3}
195. k2_funct_1(k2_funct_1(X0)) = X0 | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v2_funct_1(X0) [trivial inequality removal 194] {a:2,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
197. sK0 != sK0 | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v2_funct_1(sK0) <- (~1) [superposition 96,195] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
205. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v2_funct_1(sK0) <- (~1) [trivial inequality removal 197] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
207. ~v1_relat_1(sK0) | ~v2_funct_1(sK0) <- (~1, 3) [subsumption resolution 205,106] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
208. ~v2_funct_1(sK0) <- (~1, 3, 4) [subsumption resolution 207,111] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
209. $false <- (~1, 2, 3, 4) [subsumption resolution 208,101] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
210. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 209]
217. $false [avatar sat refutation 97,102,107,112,210] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:43,thDist:-43}
% SZS output end Proof for funct_1__t43_funct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 126
% Active clauses: 40
% Passive clauses: 53
% Final active clauses: 39
% Final passive clauses: 4
% Main loop iterations started: 49
% 
% >>> Simplifying Inferences
% Duplicate literals: 19
% Trivial inequalities: 2
% Fw subsumption resolutions: 35
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 17
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Unit resulting resolution: 1
% Forward superposition: 12
% Backward superposition: 16
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 9
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

