i: [1,38,0,2,1,1,0,0] t25_orders_1
a: [3,61,0,2,1,38]
s: 61
i: [1,37,0,2,1,1,0,0] t25_orders_1
a: [3,66,0,2,1,37]
s: 66
i: [1,39,0,4,1,1,0,0] t25_orders_1
a: [3,56,0,4,1,39]
s: 56
i: [1,46,0,4,1,0,0,255] fc6_relat_1
s: 46
i: [1,47,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 47
i: [1,40,0,6,3,0,0,1] d3_orders_1
s: 40
d: [2,68,1,4,2,1,40,40,61]
d: [2,69,1,2,1,2,21,68,66]
a: [3,73,1,2,1,69]
s: 73
i: [1,41,0,6,3,0,0,1] d3_orders_1
s: 41
d: [2,75,1,4,2,1,40,41,61]
d: [2,76,1,2,1,2,21,75,66]
a: [3,80,1,2,1,76]
s: 80
i: [1,49,0,6,2,0,0,1] dt_k2_wellord1
s: 49
i: [1,50,0,6,2,0,0,3] fc1_relat_1
s: 50
i: [1,48,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 48
i: [1,42,0,8,4,0,0,1] d3_orders_1
s: 42
d: [2,83,1,6,2,0,44,50,48]
s: 83
i: [1,44,0,8,3,0,0,2] t15_wellord1
s: 44
i: [1,45,0,8,3,0,0,2] t17_wellord1
s: 45
d: [2,84,1,6,2,0,44,50,48]
s: 84
i: [1,43,0,11,2,0,0,255] d6_wellord1
s: 43
d: [2,91,1,13,2,0,44,43,47]
d: [2,100,1,9,1,0,21,91,46]
s: 100
i: [1,51,0,12,5,0,0,2] fc22_relat_2
s: 51
i: [1,52,0,12,5,0,0,2] fc5_relat_2
s: 52
d: [2,94,1,13,2,0,44,47,43]
d: [2,101,1,9,1,0,21,94,46]
s: 101
d: [2,92,1,11,2,0,44,43,48]
s: 92
d: [2,93,1,11,2,0,44,43,48]
s: 93
d: [2,96,1,11,2,0,44,48,43]
s: 96
d: [2,97,1,11,2,0,44,48,43]
s: 97
d: [2,122,2,19,3,0,44,92,43]
d: [2,134,2,15,2,0,21,122,46]
d: [2,135,2,11,1,0,21,134,46]
s: 135
d: [2,89,1,16,5,0,40,45,42]
d: [2,90,1,12,4,0,21,89,49]
s: 90
d: [2,140,3,10,2,0,44,49,135]
d: [2,146,3,6,1,0,21,140,46]
s: 146
d: [2,105,1,20,7,0,40,51,42]
d: [2,112,1,16,6,0,21,105,83]
s: 112
d: [2,150,2,12,5,0,40,90,44]
d: [2,154,2,10,4,0,18,150]
s: 154
d: [2,170,3,0,0,4,52,73,80,66,56,154]
e: 170
d: [2,171,3,6,3,1,40,154,56]
d: [2,178,3,4,2,2,21,171,66]
d: [2,179,3,2,1,3,21,178,80]
d: [2,180,3,0,0,4,21,179,73]
e: 180
f: 180,76,69,37,38,39
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t25_orders_1
% SZS output start Proof for orders_1__t25_orders_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_orders_1(X0) => v1_orders_1(k2_wellord1(X0,X1)))) [input t25_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_orders_1(X0) => v1_orders_1(k2_wellord1(X0,X1)))) [negated conjecture 1]
4. ! [X0] : (v1_relat_1(X0) => (v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0)))) [input d3_orders_1]
6. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k2_wellord1(X0,X1))) [input dt_k2_wellord1]
14. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v1_relat_2(X1) => v1_relat_2(k2_wellord1(X1,X0)))) [input t15_wellord1]
15. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v8_relat_2(X1) => v8_relat_2(k2_wellord1(X1,X0)))) [input t17_wellord1]
16. ! [X0,X1] : (v1_relat_1(X1) => (v1_relat_2(X1) => v1_relat_2(k2_wellord1(X1,X0)))) [flattening 14]
17. ! [X0,X1] : (v1_relat_1(X1) => (v8_relat_2(X1) => v8_relat_2(k2_wellord1(X1,X0)))) [flattening 15]
19. ? [X0] : (? [X1] : (~v1_orders_1(k2_wellord1(X0,X1)) & v1_orders_1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
20. ! [X0] : ((v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
22. ! [X0,X1] : ((v1_relat_2(k2_wellord1(X1,X0)) | ~v1_relat_2(X1)) | ~v1_relat_1(X1)) [ennf transformation 16]
23. ! [X0,X1] : (v1_relat_2(k2_wellord1(X1,X0)) | ~v1_relat_2(X1) | ~v1_relat_1(X1)) [flattening 22]
24. ! [X0,X1] : ((v8_relat_2(k2_wellord1(X1,X0)) | ~v8_relat_2(X1)) | ~v1_relat_1(X1)) [ennf transformation 17]
25. ! [X0,X1] : (v8_relat_2(k2_wellord1(X1,X0)) | ~v8_relat_2(X1) | ~v1_relat_1(X1)) [flattening 24]
26. ! [X0,X1] : (v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 6]
32. ? [X0] : (? [X1] : (~v1_orders_1(k2_wellord1(X0,X1)) & v1_orders_1(X0)) & v1_relat_1(X0)) => (? [X1] : (~v1_orders_1(k2_wellord1(sK0,X1)) & v1_orders_1(sK0)) & v1_relat_1(sK0)) [choice axiom]
33. ? [X1] : (~v1_orders_1(k2_wellord1(sK0,X1)) & v1_orders_1(sK0)) => (~v1_orders_1(k2_wellord1(sK0,sK1)) & v1_orders_1(sK0)) [choice axiom]
34. (~v1_orders_1(k2_wellord1(sK0,sK1)) & v1_orders_1(sK0)) & v1_relat_1(sK0) [skolemisation 19,33,32]
35. ! [X0] : (((v1_orders_1(X0) | (~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 20]
36. ! [X0] : (((v1_orders_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 35]
37. v1_relat_1(sK0) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. v1_orders_1(sK0) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. ~v1_orders_1(k2_wellord1(sK0,sK1)) [cnf transformation 34] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. ~v1_orders_1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
41. ~v1_orders_1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. ~v8_relat_2(X0) | v1_orders_1(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. v1_relat_2(k2_wellord1(X1,X0)) | ~v1_relat_2(X1) | ~v1_relat_1(X1) [cnf transformation 23] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
45. v8_relat_2(k2_wellord1(X1,X0)) | ~v8_relat_2(X1) | ~v1_relat_1(X1) [cnf transformation 25] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
49. v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. 1 <=> v1_orders_1(k2_wellord1(sK0,sK1)) [avatar definition]
56. ~v1_orders_1(k2_wellord1(sK0,sK1)) <- (~1) [avatar component clause 54] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
57. ~1 [avatar split clause 39,54]
59. 2 <=> v1_orders_1(sK0) [avatar definition]
61. v1_orders_1(sK0) <- (2) [avatar component clause 59] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
62. 2 [avatar split clause 38,59]
64. 3 <=> v1_relat_1(sK0) [avatar definition]
66. v1_relat_1(sK0) <- (3) [avatar component clause 64] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
67. 3 [avatar split clause 37,64]
68. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 40,61] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
69. v1_relat_2(sK0) <- (2, 3) [subsumption resolution 68,66] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
71. 4 <=> v1_relat_2(sK0) [avatar definition]
73. v1_relat_2(sK0) <- (4) [avatar component clause 71] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
74. 4 | ~2 | ~3 [avatar split clause 69,64,59,71]
75. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 41,61] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
76. v8_relat_2(sK0) <- (2, 3) [subsumption resolution 75,66] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
78. 5 <=> v8_relat_2(sK0) [avatar definition]
80. v8_relat_2(sK0) <- (5) [avatar component clause 78] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
81. 5 | ~2 | ~3 [avatar split clause 76,64,59,78]
89. ~v8_relat_2(X0) | ~v1_relat_1(X0) | v1_orders_1(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_1(k2_wellord1(X0,X1)) [resolution 45,42] {a:1,w:16,wCS:6400,sine:1,thAx:0,allAx:2,thDist:-2}
90. ~v1_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_1(X0) | v1_orders_1(k2_wellord1(X0,X1)) | ~v8_relat_2(X0) [subsumption resolution 89,49] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
150. ~v1_relat_1(X0) | v1_orders_1(k2_wellord1(X0,X1)) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [resolution 90,44] {a:2,w:12,wCS:4800,sine:1,thAx:0,allAx:3,thDist:-3}
154. v1_orders_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) [duplicate literal removal 150] {a:2,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
171. ~v1_relat_1(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1) [resolution 154,56] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
178. ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1, 3) [subsumption resolution 171,66] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
179. ~v1_relat_2(sK0) <- (~1, 3, 5) [subsumption resolution 178,80] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
180. $false <- (~1, 3, 4, 5) [subsumption resolution 179,73] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
181. 1 | ~3 | ~4 | ~5 [avatar contradiction clause 180]
184. $false [avatar sat refutation 57,62,67,74,81,181] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:49,thDist:-49}
% SZS output end Proof for orders_1__t25_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 16
% Generated clauses: 125
% Active clauses: 26
% Passive clauses: 40
% Final active clauses: 26
% Final passive clauses: 9
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Duplicate literals: 17
% Fw subsumption resolutions: 23
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 38
% 
% >>> Generating Inferences
% Binary resolution: 10
% Unit resulting resolution: 1
% Forward superposition: 33
% Backward superposition: 24
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% term sharing: 0.001 s
% superposition: 0.001 s ( own 0.0000 s ) 
% other: 0.003 s ( own 0.001 s ) 

