i: [1,38,0,8,2,1,0,0] t5_card_1
a: [3,58,0,3,1,38]
s: 58
a: [3,61,0,5,1,38]
s: 61
i: [1,52,0,2,1,0,0,2] rc1_card_1
a: [3,68,0,2,1,52]
s: 68
i: [1,53,0,2,1,0,0,3] rc1_ordinal1
a: [3,73,0,2,1,53]
s: 73
i: [1,39,0,3,1,0,0,1] dt_k1_card_1
s: 39
i: [1,48,0,3,1,0,0,1] reflexivity_r2_wellord2
s: 48
i: [1,41,0,4,2,0,0,2] cc1_card_1
s: 41
i: [1,42,0,4,2,0,0,255] cc1_ordinal1
s: 42
d: [2,76,1,2,1,1,40,41,68]
a: [3,81,1,2,1,76]
s: 81
i: [1,43,0,4,2,0,0,255] cc1_ordinal1
s: 43
i: [1,40,0,6,1,0,0,1] projectivity_k1_card_1
s: 40
d: [2,83,1,2,1,1,40,42,73]
a: [3,87,1,2,1,83]
s: 87
i: [1,44,0,6,3,0,0,255] cc2_ordinal1
s: 44
d: [2,95,1,2,1,1,40,43,73]
a: [3,100,1,2,1,95]
s: 100
i: [1,49,0,6,2,0,0,1] symmetry_r2_wellord2
s: 49
i: [1,50,0,6,2,0,0,255] redefinition_r2_wellord2
s: 50
d: [2,89,2,2,1,1,40,81,42]
a: [3,93,2,2,1,89]
s: 93
i: [1,51,0,6,2,0,0,255] redefinition_r2_wellord2
s: 51
d: [2,96,2,2,1,1,40,43,81]
a: [3,105,2,2,1,96]
s: 105
i: [1,54,0,7,2,0,0,1] d2_card_1
s: 54
d: [2,114,0,4,1,0,21,54,39]
s: 114
d: [2,117,1,4,1,1,44,114,61]
a: [3,122,1,4,1,117]
s: 122
i: [1,46,0,9,3,0,0,1] d2_card_1
s: 46
d: [2,124,2,4,1,1,40,122,50]
a: [3,129,2,4,1,124]
s: 129
i: [1,47,0,9,3,0,0,1] t15_wellord2
s: 47
d: [2,136,1,9,3,1,44,46,61]
s: 136
d: [2,125,2,4,1,1,40,122,49]
a: [3,134,2,4,1,125]
s: 134
d: [2,158,3,0,0,2,52,58,114,134,47]
e: 158
r: 58
r: 136
r: 61
a: [3,57,0,3,1,38]
s: 57
d: [2,171,1,3,1,1,40,57,50]
a: [3,176,1,3,1,171]
s: 176
a: [3,62,0,5,1,38]
s: 62
d: [2,184,2,0,0,2,52,39,122,62,46]
e: 184
r: 122
r: 129
r: 134
d: [2,172,1,3,1,1,40,57,49]
a: [3,181,1,3,1,172]
s: 181
d: [2,190,2,3,1,1,40,181,50]
a: [3,195,2,3,1,190]
s: 195
a: [3,121,1,4,1,117]
s: 121
d: [2,198,2,0,0,2,52,181,114,121,47]
e: 198
i: [1,37,0,8,2,1,0,0] t5_card_1
f: 198,184,172,158,125,117,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for card_1__t5_card_1
% SZS output start Proof for card_1__t5_card_1
1. ! [X0] : ! [X1] : (r2_wellord2(X0,X1) <=> k1_card_1(X0) = k1_card_1(X1)) [input t5_card_1]
2. ~! [X0] : ! [X1] : (r2_wellord2(X0,X1) <=> k1_card_1(X0) = k1_card_1(X1)) [negated conjecture 1]
6. ! [X0] : ! [X1] : (v1_card_1(X1) => (k1_card_1(X0) = X1 <=> r2_wellord2(X0,X1))) [input d2_card_1]
7. ! [X0] : v1_card_1(k1_card_1(X0)) [input dt_k1_card_1]
13. ! [X0,X1] : (r2_wellord2(X0,X1) => r2_wellord2(X1,X0)) [input symmetry_r2_wellord2]
14. ! [X0] : ! [X1] : ! [X2] : ((r2_wellord2(X1,X2) & r2_wellord2(X0,X1)) => r2_wellord2(X0,X2)) [input t15_wellord2]
15. ~! [X0,X1] : (r2_wellord2(X0,X1) <=> k1_card_1(X0) = k1_card_1(X1)) [flattening 2]
16. ! [X0,X1] : (v1_card_1(X1) => (k1_card_1(X0) = X1 <=> r2_wellord2(X0,X1))) [flattening 6]
17. ! [X0,X1,X2] : ((r2_wellord2(X1,X2) & r2_wellord2(X0,X1)) => r2_wellord2(X0,X2)) [flattening 14]
19. ? [X0,X1] : (r2_wellord2(X0,X1) <~> k1_card_1(X0) = k1_card_1(X1)) [ennf transformation 15]
24. ! [X0,X1] : ((k1_card_1(X0) = X1 <=> r2_wellord2(X0,X1)) | ~v1_card_1(X1)) [ennf transformation 16]
25. ! [X0,X1,X2] : (r2_wellord2(X0,X2) | (~r2_wellord2(X1,X2) | ~r2_wellord2(X0,X1))) [ennf transformation 17]
26. ! [X0,X1,X2] : (r2_wellord2(X0,X2) | ~r2_wellord2(X1,X2) | ~r2_wellord2(X0,X1)) [flattening 25]
27. ! [X0,X1] : (r2_wellord2(X1,X0) | ~r2_wellord2(X0,X1)) [ennf transformation 13]
28. ? [X0,X1] : ((k1_card_1(X0) != k1_card_1(X1) | ~r2_wellord2(X0,X1)) & (k1_card_1(X0) = k1_card_1(X1) | r2_wellord2(X0,X1))) [nnf transformation 19]
29. ? [X0,X1] : ((k1_card_1(X0) != k1_card_1(X1) | ~r2_wellord2(X0,X1)) & (k1_card_1(X0) = k1_card_1(X1) | r2_wellord2(X0,X1))) => ((k1_card_1(sK0) != k1_card_1(sK1) | ~r2_wellord2(sK0,sK1)) & (k1_card_1(sK0) = k1_card_1(sK1) | r2_wellord2(sK0,sK1))) [choice axiom]
30. (k1_card_1(sK0) != k1_card_1(sK1) | ~r2_wellord2(sK0,sK1)) & (k1_card_1(sK0) = k1_card_1(sK1) | r2_wellord2(sK0,sK1)) [skolemisation 28,29]
31. ! [X0,X1] : (((k1_card_1(X0) = X1 | ~r2_wellord2(X0,X1)) & (r2_wellord2(X0,X1) | k1_card_1(X0) != X1)) | ~v1_card_1(X1)) [nnf transformation 24]
37. k1_card_1(sK0) = k1_card_1(sK1) | r2_wellord2(sK0,sK1) [cnf transformation 30] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. k1_card_1(sK0) != k1_card_1(sK1) | ~r2_wellord2(sK0,sK1) [cnf transformation 30] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. v1_card_1(k1_card_1(X0)) [cnf transformation 7] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. r2_wellord2(X0,X1) | k1_card_1(X0) != X1 | ~v1_card_1(X1) [cnf transformation 31] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:1,thDist:-1}
46. k1_card_1(X0) = X1 | ~r2_wellord2(X0,X1) | ~v1_card_1(X1) [cnf transformation 31] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. ~r2_wellord2(X1,X2) | r2_wellord2(X0,X2) | ~r2_wellord2(X0,X1) [cnf transformation 26] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. ~r2_wellord2(X0,X1) | r2_wellord2(X1,X0) [cnf transformation 27] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. r2_wellord2(X0,k1_card_1(X0)) | ~v1_card_1(k1_card_1(X0)) [equality resolution 45] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
56. 1 <=> r2_wellord2(sK0,sK1) [avatar definition]
57. r2_wellord2(sK0,sK1) <- (1) [avatar component clause 56] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
58. ~r2_wellord2(sK0,sK1) <- (~1) [avatar component clause 56] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 2 <=> k1_card_1(sK0) = k1_card_1(sK1) [avatar definition]
61. k1_card_1(sK0) = k1_card_1(sK1) <- (2) [avatar component clause 60] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. k1_card_1(sK0) != k1_card_1(sK1) <- (~2) [avatar component clause 60] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. ~1 | ~2 [avatar split clause 38,60,56]
64. 1 | 2 [avatar split clause 37,60,56]
114. r2_wellord2(X0,k1_card_1(X0)) [subsumption resolution 54,39] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
117. r2_wellord2(sK1,k1_card_1(sK0)) <- (2) [superposition 114,61] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
120. 10 <=> r2_wellord2(sK1,k1_card_1(sK0)) [avatar definition]
121. ~r2_wellord2(sK1,k1_card_1(sK0)) <- (~10) [avatar component clause 120] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
122. r2_wellord2(sK1,k1_card_1(sK0)) <- (10) [avatar component clause 120] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
123. 10 | ~2 [avatar split clause 117,60,120]
125. r2_wellord2(k1_card_1(sK0),sK1) <- (10) [resolution 122,49] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
132. 12 <=> r2_wellord2(k1_card_1(sK0),sK1) [avatar definition]
134. r2_wellord2(k1_card_1(sK0),sK1) <- (12) [avatar component clause 132] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
135. 12 | ~10 [avatar split clause 125,120,132]
158. $false <- (~1, 12) [unit resulting resolution 58,114,134,47] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
164. 1 | ~12 [avatar contradiction clause 158]
172. r2_wellord2(sK1,sK0) <- (1) [resolution 57,49] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
179. 15 <=> r2_wellord2(sK1,sK0) [avatar definition]
181. r2_wellord2(sK1,sK0) <- (15) [avatar component clause 179] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
182. 15 | ~1 [avatar split clause 172,56,179]
184. $false <- (~2, 10) [unit resulting resolution 39,122,62,46] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
186. 2 | ~10 [avatar contradiction clause 184]
198. $false <- (~10, 15) [unit resulting resolution 181,114,121,47] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
200. 10 | ~15 [avatar contradiction clause 198]
201. $false [avatar sat refutation 63,64,123,135,164,182,186,200] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:118,thDist:-118}
% SZS output end Proof for card_1__t5_card_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 105
% Active clauses: 32
% Passive clauses: 59
% Final active clauses: 26
% Final passive clauses: 15
% Main loop iterations started: 33
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 15
% 
% >>> Generating Inferences
% Binary resolution: 29
% Unit resulting resolution: 3
% Forward superposition: 12
% Backward superposition: 13
% Self superposition: 3
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 2
% SAT solver binary clauses: 19
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

