i: [1,241,0,2,1,1,0,0] t184_xreal_1
a: [3,387,0,2,1,241]
s: 387
i: [1,240,0,2,1,1,0,0] t184_xreal_1
a: [3,392,0,2,1,240]
s: 392
i: [1,243,0,3,1,1,0,0] t184_xreal_1
a: [3,377,0,3,1,243]
s: 377
i: [1,242,0,3,1,1,0,0] t184_xreal_1
a: [3,382,0,3,1,242]
s: 382
i: [1,244,0,5,1,1,0,0] t184_xreal_1
a: [3,372,0,5,1,244]
s: 372
i: [1,245,0,2,1,0,0,255] fc1_numbers
a: [3,397,0,2,1,245]
s: 397
i: [1,246,0,2,1,0,0,1] spc1_boole
a: [3,402,0,2,1,246]
s: 402
i: [1,247,0,2,1,0,0,255] fc1_xboole_0
a: [3,407,0,2,1,247]
s: 407
i: [1,248,0,2,1,0,0,255] spc0_boole
a: [3,412,0,2,1,248]
s: 412
i: [1,260,0,2,1,0,0,255] fc6_ordinal1
a: [3,467,0,2,1,260]
s: 467
i: [1,259,0,2,1,0,0,255] fc6_ordinal1
a: [3,472,0,2,1,259]
s: 472
i: [1,261,0,2,1,0,0,255] spc1_numerals
a: [3,492,0,2,1,261]
s: 492
i: [1,336,0,2,1,0,0,2] rc2_xboole_0
a: [3,519,0,2,1,336]
s: 519
i: [1,337,0,2,1,0,0,255] rc4_ordinal1
a: [3,524,0,2,1,337]
s: 524
i: [1,338,0,2,1,0,0,255] rc1_ordinal1
a: [3,529,0,2,1,338]
s: 529
i: [1,339,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,534,0,2,1,339]
s: 534
i: [1,340,0,2,1,0,0,1] rc1_xreal_0
a: [3,539,0,2,1,340]
s: 539
i: [1,341,0,2,1,0,0,2] rc1_xxreal_0
a: [3,544,0,2,1,341]
s: 544
i: [1,342,0,2,1,0,0,2] rc1_xboole_0
a: [3,549,0,2,1,342]
s: 549
i: [1,346,0,2,1,0,0,255] rc2_ordinal1
a: [3,554,0,2,1,346]
s: 554
i: [1,345,0,2,1,0,0,255] rc2_ordinal1
a: [3,559,0,2,1,345]
s: 559
i: [1,344,0,2,1,0,0,255] rc2_ordinal1
a: [3,564,0,2,1,344]
s: 564
i: [1,343,0,2,1,0,0,255] rc2_ordinal1
a: [3,569,0,2,1,343]
s: 569
i: [1,348,0,2,1,0,0,255] rc5_ordinal1
a: [3,574,0,2,1,348]
s: 574
i: [1,347,0,2,1,0,0,255] rc5_ordinal1
a: [3,579,0,2,1,347]
s: 579
i: [1,350,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,584,0,2,1,350]
s: 584
i: [1,349,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,589,0,2,1,349]
s: 589
i: [1,354,0,2,1,0,0,2] rc3_xreal_0
a: [3,594,0,2,1,354]
s: 594
i: [1,353,0,2,1,0,0,2] rc3_xreal_0
a: [3,599,0,2,1,353]
s: 599
i: [1,352,0,2,1,0,0,2] rc3_xreal_0
a: [3,604,0,2,1,352]
s: 604
i: [1,351,0,2,1,0,0,2] rc3_xreal_0
a: [3,609,0,2,1,351]
s: 609
i: [1,358,0,2,1,0,0,2] rc2_xreal_0
a: [3,614,0,2,1,358]
s: 614
i: [1,357,0,2,1,0,0,2] rc2_xreal_0
a: [3,619,0,2,1,357]
s: 619
i: [1,356,0,2,1,0,0,2] rc2_xreal_0
a: [3,624,0,2,1,356]
s: 624
i: [1,355,0,2,1,0,0,2] rc2_xreal_0
a: [3,629,0,2,1,355]
s: 629
i: [1,360,0,2,1,0,0,2] rc3_xxreal_0
a: [3,634,0,2,1,360]
s: 634
i: [1,359,0,2,1,0,0,2] rc3_xxreal_0
a: [3,639,0,2,1,359]
s: 639
i: [1,362,0,2,1,0,0,2] rc2_xxreal_0
a: [3,644,0,2,1,362]
s: 644
i: [1,361,0,2,1,0,0,2] rc2_xxreal_0
a: [3,649,0,2,1,361]
s: 649
i: [1,366,0,2,1,0,0,2] rc4_xreal_0
a: [3,654,0,2,1,366]
s: 654
i: [1,365,0,2,1,0,0,2] rc4_xreal_0
a: [3,659,0,2,1,365]
s: 659
i: [1,364,0,2,1,0,0,2] rc4_xreal_0
a: [3,664,0,2,1,364]
s: 664
i: [1,363,0,2,1,0,0,2] rc4_xreal_0
a: [3,669,0,2,1,363]
s: 669
i: [1,368,0,2,1,0,0,2] rc4_xxreal_0
a: [3,674,0,2,1,368]
s: 674
i: [1,367,0,2,1,0,0,2] rc4_xxreal_0
a: [3,679,0,2,1,367]
s: 679
i: [1,321,0,3,1,0,0,255] reflexivity_r1_tarski
s: 321
i: [1,249,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r1_r0
a: [3,417,0,3,1,249]
s: 417
i: [1,250,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,422,0,3,1,250]
s: 422
i: [1,251,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,427,0,3,1,251]
s: 427
r: 407
d: [2,681,0,2,1,2,23,407,427]
a: [3,685,0,2,1,681]
s: 685
i: [1,252,0,3,1,0,0,255] t1_numerals
a: [3,432,0,3,1,252]
s: 432
i: [1,253,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r0_r0
a: [3,437,0,3,1,253]
s: 437
i: [1,254,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r0_r1
a: [3,442,0,3,1,254]
s: 442
i: [1,255,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,447,0,3,1,255]
s: 447
i: [1,264,0,3,1,0,0,255] spc1_numerals
a: [3,477,0,3,1,264]
s: 477
i: [1,263,0,3,1,0,0,255] spc1_numerals
a: [3,482,0,3,1,263]
s: 482
i: [1,267,0,3,1,0,0,255] spc0_numerals
a: [3,497,0,3,1,267]
s: 497
i: [1,266,0,3,1,0,0,255] spc0_numerals
a: [3,502,0,3,1,266]
s: 502
d: [2,687,0,3,1,2,22,432,427]
a: [3,691,0,3,1,687]
s: 691
d: [2,693,0,3,1,2,22,482,422]
a: [3,697,0,3,1,693]
s: 697
d: [2,699,0,3,1,2,22,502,422]
a: [3,703,0,3,1,699]
s: 703
i: [1,269,0,4,2,0,0,255] cc6_ordinal1
s: 269
i: [1,270,0,4,2,0,0,255] cc2_xcmplx_0
s: 270
d: [2,705,1,2,1,1,40,269,524]
a: [3,710,1,2,1,705]
s: 710
i: [1,271,0,4,2,0,0,255] cc2_xreal_0
s: 271
i: [1,272,0,4,2,0,0,255] cc2_xxreal_0
s: 272
d: [2,706,1,2,1,1,40,269,574]
a: [3,715,1,2,1,706]
s: 715
i: [1,273,0,4,2,0,0,255] cc1_ordinal1
s: 273
d: [2,717,1,2,1,1,40,270,524]
a: [3,722,1,2,1,717]
s: 722
i: [1,274,0,4,2,0,0,255] cc1_ordinal1
s: 274
i: [1,279,0,4,2,0,0,2] cc3_xreal_0
s: 279
d: [2,803,1,2,1,1,40,279,392]
a: [3,812,1,2,1,803]
s: 812
i: [1,280,0,4,2,0,0,2] cc4_xreal_0
s: 280
d: [2,804,1,2,1,1,40,279,387]
a: [3,817,1,2,1,804]
s: 817
i: [1,290,0,4,2,0,0,255] cc9_ordinal1
s: 290
i: [1,291,0,4,2,0,0,255] cc7_ordinal1
s: 291
d: [2,824,1,2,1,1,40,280,392]
a: [3,833,1,2,1,824]
s: 833
i: [1,292,0,4,2,0,0,255] cc3_ordinal1
s: 292
d: [2,825,1,2,1,1,40,280,387]
a: [3,838,1,2,1,825]
s: 838
i: [1,320,0,4,1,0,0,255] existence_m1_subset_1
s: 320
i: [1,256,0,4,1,0,0,255] dt_k5_numbers
a: [3,452,0,4,1,256]
s: 452
d: [2,718,1,2,1,1,40,270,574]
a: [3,727,1,2,1,718]
s: 727
i: [1,257,0,4,1,0,0,255] dt_k6_numbers
a: [3,457,0,4,1,257]
s: 457
d: [2,729,1,2,1,1,40,271,524]
a: [3,734,1,2,1,729]
s: 734
i: [1,262,0,4,1,0,0,255] spc1_numerals
a: [3,487,0,4,1,262]
s: 487
i: [1,265,0,4,1,0,0,255] spc0_numerals
a: [3,507,0,4,1,265]
s: 507
d: [2,730,1,2,1,1,40,271,574]
a: [3,739,1,2,1,730]
s: 739
d: [2,935,0,4,1,2,22,452,422]
a: [3,939,0,4,1,935]
s: 939
d: [2,741,1,2,1,1,40,272,524]
a: [3,746,1,2,1,741]
s: 746
d: [2,941,0,4,1,2,22,457,422]
a: [3,945,0,4,1,941]
s: 945
d: [2,950,0,4,1,2,22,487,422]
a: [3,954,0,4,1,950]
s: 954
d: [2,742,1,2,1,1,40,272,574]
a: [3,751,1,2,1,742]
s: 751
d: [2,956,0,4,1,2,22,507,422]
a: [3,960,0,4,1,956]
s: 960
d: [2,753,1,2,1,1,40,273,467]
a: [3,761,1,2,1,753]
s: 761
i: [1,293,0,5,2,0,0,255] t6_boole
s: 293
i: [1,294,0,5,2,0,0,255] cc8_ordinal1
s: 294
d: [2,755,1,2,1,1,40,273,529]
a: [3,771,1,2,1,755]
s: 771
i: [1,295,0,5,2,0,0,255] cc1_xcmplx_0
s: 295
d: [2,778,1,2,1,1,40,274,467]
a: [3,786,1,2,1,778]
s: 786
i: [1,296,0,5,2,0,0,255] cc1_xreal_0
s: 296
i: [1,317,0,5,2,0,0,255] t7_boole
s: 317
d: [2,780,1,2,1,1,40,274,529]
a: [3,796,1,2,1,780]
s: 796
i: [1,258,0,5,1,0,0,2] rqRealDiv__k7_xcmplx_0__r1_r1_r1
a: [3,462,0,5,1,258]
s: 462
d: [2,805,1,2,1,1,40,279,539]
a: [3,822,1,2,1,805]
s: 822
i: [1,334,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,514,0,5,2,334]
s: 514
d: [2,965,0,5,2,1,22,293,427]
s: 965
d: [2,1022,1,5,2,2,44,377,965]
s: 1022
i: [1,301,0,6,3,0,0,255] cc2_ordinal1
s: 301
d: [2,826,1,2,1,1,40,280,539]
a: [3,843,1,2,1,826]
s: 843
i: [1,304,0,6,3,0,0,2] cc5_xxreal_0
s: 304
i: [1,306,0,6,3,0,0,2] cc5_xxreal_0
s: 306
d: [2,845,1,2,1,1,40,290,685]
a: [3,853,1,2,1,845]
s: 853
i: [1,307,0,6,3,0,0,2] cc3_xxreal_0
s: 307
d: [2,846,1,2,1,1,40,290,412]
a: [3,858,1,2,1,846]
s: 858
i: [1,309,0,6,3,0,0,2] cc3_xxreal_0
s: 309
i: [1,311,0,6,3,0,0,2] cc7_xxreal_0
s: 311
d: [2,847,1,2,1,1,40,290,549]
a: [3,863,1,2,1,847]
s: 863
i: [1,312,0,6,3,0,0,2] cc7_xxreal_0
s: 312
d: [2,848,1,2,1,1,40,290,669]
a: [3,868,1,2,1,848]
s: 868
i: [1,313,0,6,2,0,0,255] t1_subset
s: 313
i: [1,322,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 322
d: [2,849,1,2,1,1,40,290,679]
a: [3,873,1,2,1,849]
s: 873
i: [1,275,0,7,3,0,0,255] cc5_ordinal1
s: 275
d: [2,875,1,2,1,1,40,291,685]
a: [3,883,1,2,1,875]
s: 883
i: [1,276,0,7,2,0,0,1] t5_arithm
s: 276
i: [1,277,0,7,2,0,0,2] t6_arithm
s: 277
d: [2,876,1,2,1,1,40,291,412]
a: [3,888,1,2,1,876]
s: 888
i: [1,315,0,7,2,0,0,255] t3_subset
s: 315
d: [2,877,1,2,1,1,40,291,549]
a: [3,893,1,2,1,877]
s: 893
i: [1,316,0,7,3,0,0,2] t8_boole
s: 316
i: [1,268,0,8,3,0,0,255] cc10_ordinal1
s: 268
d: [2,878,1,2,1,1,40,291,669]
a: [3,898,1,2,1,878]
s: 898
i: [1,298,0,8,4,0,0,2] cc4_xxreal_0
s: 298
d: [2,1165,2,6,3,1,40,298,833]
a: [3,1180,2,2,1,1165]
s: 1180
r: 1180
i: [1,300,0,8,4,0,0,2] cc6_xxreal_0
s: 300
i: [1,302,0,8,4,0,0,2] cc8_xxreal_0
s: 302
a: [3,1179,2,2,1,1165]
s: 1179
i: [1,314,0,8,3,0,0,255] t2_subset
s: 314
a: [3,1184,2,2,1,1165]
s: 1184
i: [1,332,0,8,3,0,0,1] fc5_xcmplx_0
s: 332
i: [1,333,0,8,3,0,0,1] fc8_xreal_0
s: 333
d: [2,1166,2,6,3,1,40,298,838]
a: [3,1193,2,2,1,1166]
s: 1193
r: 1193
i: [1,319,0,9,3,0,0,255] t5_subset
s: 319
a: [3,1192,2,2,1,1166]
s: 1192
i: [1,278,0,10,3,0,0,1] t60_xcmplx_1
s: 278
i: [1,318,0,10,3,0,0,255] t4_subset
s: 318
a: [3,1197,2,2,1,1166]
s: 1197
i: [1,335,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 335
d: [2,1370,1,11,3,1,40,335,372]
a: [3,1377,1,2,1,1370]
s: 1377
i: [1,283,0,11,5,0,0,1] t2_real
s: 283
i: [1,284,0,11,5,0,0,1] t1_real
s: 284
d: [2,1429,1,8,4,1,40,284,377]
d: [2,1436,1,6,3,2,21,1429,1197]
d: [2,1437,1,4,2,3,21,1436,387]
a: [3,1441,1,2,1,1437]
s: 1441
r: 1197
r: 1441
r: 1184
i: [1,285,0,11,5,0,0,1] t8_real
s: 285
a: [3,1440,1,2,1,1437]
s: 1440
i: [1,286,0,11,5,0,0,1] t7_real
s: 286
i: [1,287,0,11,5,0,0,1] t3_real
s: 287
a: [3,1183,2,2,1,1165]
s: 1183
i: [1,288,0,11,5,0,0,1] t4_real
s: 288
a: [3,1188,2,2,1,1165]
s: 1188
i: [1,323,0,12,5,0,0,1] fc33_xreal_0
s: 323
i: [1,324,0,12,5,0,0,1] fc31_xreal_0
s: 324
a: [3,1196,2,2,1,1166]
s: 1196
i: [1,325,0,12,5,0,0,1] fc32_xreal_0
s: 325
a: [3,1201,2,2,1,1166]
s: 1201
i: [1,326,0,12,5,0,0,1] fc34_xreal_0
s: 326
i: [1,331,0,12,5,0,0,1] fc9_xcmplx_0
s: 331
d: [2,1393,1,8,4,1,40,283,377]
d: [2,1412,1,6,3,2,21,1393,387]
a: [3,1416,1,2,1,1412]
s: 1416
i: [1,281,0,13,6,0,0,1] t5_real
s: 281
d: [2,1463,1,12,4,1,40,285,372]
a: [3,1472,1,4,1,1463]
s: 1472
d: [2,1556,2,0,0,3,52,387,392,1472,333]
e: 1556
d: [2,1557,2,4,2,1,40,1472,333]
d: [2,1559,2,2,1,2,21,1557,392]
d: [2,1560,2,0,0,3,21,1559,387]
e: 1560
r: 1472
i: [1,282,0,13,6,0,0,1] t6_real
s: 282
i: [1,330,0,14,4,0,0,255] existence_m2_subset_1
s: 330
a: [3,1471,1,4,1,1463]
s: 1471
i: [1,327,0,15,5,0,0,255] dt_m2_subset_1
s: 327
a: [3,1380,1,4,1,1370]
s: 1380
i: [1,328,0,15,5,0,0,255] redefinition_m2_subset_1
s: 328
i: [1,329,0,15,5,0,0,255] redefinition_m2_subset_1
s: 329
a: [3,1476,1,4,1,1463]
s: 1476
i: [1,289,0,19,6,0,0,1] l102_xreal_1
s: 289
d: [2,1478,2,4,2,2,44,1440,965]
s: 1478
r: 307
d: [2,879,1,2,1,1,40,291,679]
a: [3,903,1,2,1,879]
s: 903
d: [2,905,1,2,1,1,40,292,685]
a: [3,913,1,2,1,905]
s: 913
d: [2,1535,2,4,2,2,44,1416,965]
s: 1535
r: 304
d: [2,906,1,2,1,1,40,292,412]
a: [3,918,1,2,1,906]
s: 918
d: [2,1577,2,4,1,1,40,1471,279]
a: [3,1582,2,4,1,1577]
s: 1582
d: [2,907,1,2,1,1,40,292,549]
a: [3,923,1,2,1,907]
s: 923
d: [2,908,1,2,1,1,40,292,669]
a: [3,928,1,2,1,908]
s: 928
d: [2,1597,2,12,3,1,40,1380,298]
a: [3,1600,2,4,1,1597]
s: 1600
d: [2,909,1,2,1,1,40,292,679]
a: [3,933,1,2,1,909]
s: 933
d: [2,967,1,2,1,1,40,294,697]
a: [3,974,1,2,1,967]
s: 974
d: [2,1769,2,0,0,2,52,1377,974,272]
e: 1769
d: [2,1770,2,2,1,1,40,974,272]
d: [2,1775,2,0,0,2,21,1770,1377]
e: 1775
r: 1377
a: [3,1376,1,2,1,1370]
s: 1376
a: [3,1385,1,5,1,1370]
s: 1385
a: [3,1604,2,4,1,1597]
s: 1604
d: [2,982,1,2,1,1,40,295,477]
a: [3,988,1,2,1,982]
s: 988
d: [2,983,1,2,1,1,40,295,497]
a: [3,993,1,2,1,983]
s: 993
d: [2,1000,1,2,1,1,40,296,477]
a: [3,1006,1,2,1,1000]
s: 1006
d: [2,1001,1,2,1,1,40,296,497]
a: [3,1011,1,2,1,1001]
s: 1011
d: [2,1026,1,4,2,1,40,304,599]
d: [2,1028,1,2,1,2,21,1026,604]
a: [3,1032,1,2,1,1028]
s: 1032
d: [2,1027,1,4,2,1,40,304,634]
d: [2,1034,1,2,1,2,21,1027,639]
a: [3,1038,1,2,1,1034]
s: 1038
d: [2,1040,1,4,2,1,40,306,599]
d: [2,1042,1,2,1,2,21,1040,604]
a: [3,1046,1,2,1,1042]
s: 1046
d: [2,1041,1,4,2,1,40,306,634]
d: [2,1048,1,2,1,2,21,1041,639]
a: [3,1052,1,2,1,1048]
s: 1052
d: [2,1056,1,4,2,1,40,307,619]
d: [2,1058,1,2,1,2,21,1056,624]
a: [3,1062,1,2,1,1058]
s: 1062
d: [2,1057,1,4,2,1,40,307,644]
d: [2,1064,1,2,1,2,21,1057,649]
a: [3,1068,1,2,1,1064]
s: 1068
d: [2,1072,1,4,2,1,40,275,477]
a: [3,1084,1,2,1,1072]
s: 1084
d: [2,1148,1,4,2,1,40,268,939]
a: [3,1154,1,2,1,1148]
s: 1154
d: [2,1169,1,6,3,1,40,298,544]
a: [3,1232,1,2,1,1169]
s: 1232
r: 1232
a: [3,1231,1,2,1,1169]
s: 1231
a: [3,1236,1,2,1,1169]
s: 1236
d: [2,1391,1,8,4,1,40,283,442]
a: [3,1407,1,2,1,1391]
s: 1407
d: [2,754,2,2,1,1,40,273,710]
a: [3,766,2,2,1,754]
s: 766
d: [2,969,1,3,1,0,40,294,320]
a: [3,980,1,3,1,969]
s: 980
d: [2,984,1,3,1,0,40,295,320]
a: [3,998,1,3,1,984]
s: 998
d: [2,757,2,2,1,1,40,273,715]
a: [3,776,2,2,1,757]
s: 776
d: [2,1002,1,3,1,0,40,296,320]
a: [3,1016,1,3,1,1002]
s: 1016
d: [2,779,2,2,1,1,40,274,710]
a: [3,791,2,2,1,779]
s: 791
d: [2,1273,1,5,2,1,40,314,477]
d: [2,1287,1,3,1,2,21,1273,397]
a: [3,1291,1,3,1,1287]
s: 1291
d: [2,1274,1,5,2,1,40,314,497]
d: [2,1293,1,3,1,2,21,1274,397]
a: [3,1297,1,3,1,1293]
s: 1297
d: [2,782,2,2,1,1,40,274,715]
a: [3,801,2,2,1,782]
s: 801
d: [2,1275,1,5,2,1,40,314,691]
d: [2,1299,1,3,1,2,21,1275,472]
a: [3,1303,1,3,1,1299]
s: 1303
d: [2,1093,2,2,1,1,40,883,272]
a: [3,1101,2,2,1,1093]
s: 1101
d: [2,1276,1,5,2,1,40,314,697]
d: [2,1305,1,3,1,2,21,1276,472]
a: [3,1309,1,3,1,1305]
s: 1309
d: [2,1277,1,5,2,1,40,314,703]
d: [2,1311,1,3,1,2,21,1277,472]
a: [3,1315,1,3,1,1311]
s: 1315
d: [2,1094,2,2,1,1,40,883,271]
a: [3,1106,2,2,1,1094]
s: 1106
d: [2,1584,1,11,4,1,40,327,945]
d: [2,1589,1,7,3,2,21,1584,939]
d: [2,1590,1,5,2,3,21,1589,472]
d: [2,1591,1,3,1,4,21,1590,397]
a: [3,1595,1,3,1,1591]
s: 1595
d: [2,1095,2,2,1,1,40,883,270]
a: [3,1111,2,2,1,1095]
s: 1111
d: [2,1626,1,17,6,1,44,289,462]
d: [2,1637,1,15,5,1,18,1626]
a: [3,1649,1,3,1,1637]
s: 1649
r: 1649
a: [3,1648,1,3,1,1637]
s: 1648
d: [2,1927,3,0,0,6,52,1006,1106,1407,402,1440,1648,282]
e: 1927
d: [2,1928,3,0,0,5,52,1006,1106,1440,492,1648,284]
e: 1928
d: [2,1929,2,10,5,1,40,1648,282]
d: [2,1936,2,8,4,2,21,1929,1440]
d: [2,1937,2,6,3,3,21,1936,402]
d: [2,1938,2,4,2,4,21,1937,1407]
d: [2,1939,2,2,1,5,21,1938,1106]
d: [2,1940,2,0,0,6,21,1939,1006]
e: 1940
d: [2,1931,2,8,4,1,40,1648,284]
d: [2,1942,2,6,3,2,21,1931,492]
d: [2,1943,2,4,2,3,21,1942,1440]
d: [2,1944,2,2,1,4,21,1943,1106]
d: [2,1945,2,0,0,5,21,1944,1006]
e: 1945
r: 1648
d: [2,1918,2,7,3,1,40,1649,335]
d: [2,1920,2,5,2,2,21,1918,1376]
d: [2,1921,2,3,1,3,21,1920,1101]
a: [3,1925,2,3,1,1921]
s: 1925
s: 1649
d: [2,1951,3,5,2,2,44,1925,965]
s: 1951
d: [2,1271,1,7,2,1,40,314,939]
a: [3,1281,1,4,1,1271]
s: 1281
a: [3,1284,1,3,1,1271]
s: 1284
d: [2,1115,2,2,1,1,40,888,272]
a: [3,1122,2,2,1,1115]
s: 1122
d: [2,1020,1,5,2,2,44,691,965]
s: 1020
d: [2,1128,2,2,1,1,40,893,272]
a: [3,1135,2,2,1,1128]
s: 1135
d: [2,1070,1,5,2,0,40,275,320]
s: 1070
d: [2,1019,1,6,2,2,44,945,965]
s: 1019
d: [2,1129,2,2,1,1,40,893,271]
a: [3,1140,2,2,1,1129]
s: 1140
d: [2,1150,1,6,2,0,40,268,320]
s: 1150
d: [2,1130,2,2,1,1,40,893,270]
a: [3,1145,2,2,1,1130]
s: 1145
d: [2,1270,1,6,2,0,40,314,320]
s: 1270
d: [2,1366,1,6,2,1,40,318,939]
s: 1366
d: [2,1167,2,6,3,1,40,298,746]
a: [3,1206,2,2,1,1167]
s: 1206
r: 1206
d: [2,1018,1,7,3,1,44,965,965]
s: 1018
a: [3,1205,2,2,1,1167]
s: 1205
d: [2,1149,1,7,3,0,40,268,315]
s: 1149
d: [2,1350,1,7,2,0,40,319,320]
s: 1350
a: [3,1210,2,2,1,1167]
s: 1210
d: [2,1612,1,11,4,1,40,329,939]
d: [2,1615,1,9,3,2,21,1612,472]
d: [2,1616,1,7,2,3,21,1615,397]
s: 1616
d: [2,1168,2,6,3,1,40,298,843]
a: [3,1219,2,2,1,1168]
s: 1219
r: 1219
d: [2,1127,1,8,3,0,40,315,275]
s: 1127
d: [2,1324,1,8,3,0,40,333,280]
s: 1324
a: [3,1218,2,2,1,1168]
s: 1218
d: [2,1325,1,8,3,0,40,333,279]
s: 1325
a: [3,1223,2,2,1,1168]
s: 1223
d: [2,1349,1,8,3,0,40,319,315]
s: 1349
d: [2,1368,1,8,2,0,40,318,320]
s: 1368
d: [2,1170,2,6,3,1,40,298,751]
d: [2,1242,2,4,2,2,21,1170,579]
a: [3,1246,2,2,1,1242]
s: 1246
d: [2,1537,1,18,9,0,40,281,285]
d: [2,1547,1,14,7,0,18,1537]
a: [3,1551,1,8,4,1547]
s: 1551
r: 1551
d: [2,1688,2,2,1,1,40,903,271]
a: [3,1694,2,2,1,1688]
s: 1694
a: [3,1554,1,6,3,1547]
s: 1554
d: [2,2081,3,0,0,4,52,1416,1440,1106,1554]
e: 2081
d: [2,2082,3,4,2,2,40,1554,1106]
d: [2,2099,3,2,1,3,21,2082,1440]
d: [2,2100,3,0,0,4,21,2099,1416]
e: 2100
r: 1535
r: 1416
r: 1554
s: 304
d: [2,1511,1,12,6,0,44,326,276]
d: [2,1516,1,10,5,1,21,1511,1440]
d: [2,1517,1,8,4,1,21,1516,279]
a: [3,1520,1,4,2,1517]
s: 1520
r: 326
r: 325
r: 324
r: 285
r: 284
r: 282
r: 281
d: [2,2140,3,0,0,3,52,1440,1106,1520]
e: 2140
d: [2,2141,3,0,0,3,52,1183,392,1520]
e: 2141
d: [2,2142,3,0,0,3,52,1196,387,1520]
e: 2142
d: [2,2143,2,0,0,3,52,1046,594,1520]
e: 2143
d: [2,2144,3,2,1,2,40,1520,1106]
d: [2,2164,3,0,0,3,21,2144,1440]
e: 2164
d: [2,2149,2,2,1,2,40,1520,392]
d: [2,2166,2,0,0,3,21,2149,1183]
e: 2166
d: [2,2150,2,2,1,2,40,1520,387]
d: [2,2168,2,0,0,3,21,2150,1196]
e: 2168
d: [2,2156,2,2,1,2,40,1520,594]
d: [2,2170,2,0,0,3,21,2156,1046]
e: 2170
r: 1520
s: 284
s: 1416
s: 285
s: 324
d: [2,2191,2,4,2,2,44,1416,965]
s: 2191
r: 304
s: 325
d: [2,1689,2,2,1,1,40,903,270]
a: [3,1699,2,2,1,1689]
s: 1699
s: 326
s: 281
d: [2,1702,2,2,1,1,40,913,274]
a: [3,1708,2,2,1,1702]
s: 1708
s: 282
d: [2,1703,2,2,1,1,40,913,273]
a: [3,1713,2,2,1,1703]
s: 1713
s: 1551
d: [2,1114,1,9,3,1,44,276,965]
s: 1114
d: [2,1721,2,2,1,1,40,918,274]
a: [3,1726,2,2,1,1721]
s: 1726
d: [2,1367,1,9,3,0,40,318,315]
s: 1367
d: [2,1722,2,2,1,1,40,918,273]
a: [3,1731,2,2,1,1722]
s: 1731
d: [2,1495,1,15,7,0,44,324,278]
d: [2,1496,1,13,6,0,18,1495]
d: [2,1498,1,11,5,1,21,1496,492]
d: [2,1499,1,9,4,1,21,1498,279]
s: 1499
d: [2,1503,1,15,7,0,44,325,278]
d: [2,1504,1,13,6,0,18,1503]
d: [2,1506,1,11,5,1,21,1504,492]
d: [2,1507,1,9,4,1,21,1506,279]
s: 1507
d: [2,1733,2,2,1,1,40,923,274]
a: [3,1738,2,2,1,1733]
s: 1738
d: [2,2203,1,15,7,0,44,324,278]
d: [2,2204,1,13,6,0,18,2203]
d: [2,2206,1,11,5,1,21,2204,492]
d: [2,2207,1,9,4,1,21,2206,279]
s: 2207
d: [2,1734,2,2,1,1,40,923,273]
a: [3,1743,2,2,1,1734]
s: 1743
d: [2,2216,1,15,7,0,44,325,278]
d: [2,2217,1,13,6,0,18,2216]
d: [2,2219,1,11,5,1,21,2217,492]
d: [2,2220,1,9,4,1,21,2219,279]
s: 2220
d: [2,1272,1,10,3,0,40,314,315]
s: 1272
d: [2,1745,2,2,1,1,40,928,274]
a: [3,1750,2,2,1,1745]
s: 1750
a: [3,1652,1,10,3,1637]
s: 1652
d: [2,2349,2,10,4,1,44,1652,276]
d: [2,2362,2,7,3,2,21,2349,1649]
d: [2,2363,2,5,2,3,21,2362,1106]
d: [2,2364,2,2,1,4,21,2363,1925]
d: [2,2365,2,0,0,5,21,2364,988]
e: 2365
r: 1652
d: [2,1746,2,2,1,1,40,928,273]
a: [3,1755,2,2,1,1746]
s: 1755
d: [2,1630,1,17,6,1,44,289,462]
d: [2,1635,1,15,5,1,18,1630]
a: [3,1667,1,10,3,1635]
s: 1667
d: [2,1388,1,15,7,0,40,283,335]
d: [2,1398,1,13,6,0,21,1388,280]
d: [2,1399,1,11,5,0,21,1398,280]
s: 1399
d: [2,1757,2,2,1,1,40,933,274]
a: [3,1762,2,2,1,1757]
s: 1762
d: [2,1424,1,15,7,0,40,284,335]
d: [2,1434,1,13,6,0,21,1424,280]
d: [2,1435,1,11,5,0,21,1434,280]
s: 1435
d: [2,1758,2,2,1,1,40,933,273]
a: [3,1767,2,2,1,1758]
s: 1767
d: [2,2173,1,15,7,0,40,284,335]
d: [2,2187,1,13,6,0,21,2173,280]
d: [2,2188,1,11,5,0,21,2187,280]
s: 2188
d: [2,1628,1,19,7,0,44,289,277]
d: [2,1659,1,17,6,0,21,1628,279]
a: [3,1662,1,12,4,1659]
s: 1662
r: 1662
d: [2,1793,2,2,1,1,40,1084,274]
a: [3,1798,2,2,1,1793]
s: 1798
d: [2,1632,1,19,7,0,44,289,277]
d: [2,1674,1,17,6,0,21,1632,279]
a: [3,1677,1,12,4,1674]
s: 1677
r: 1677
d: [2,1794,2,2,1,1,40,1084,273]
a: [3,1803,2,2,1,1794]
s: 1803
d: [2,1587,1,21,7,0,40,327,330]
d: [2,1588,1,13,4,0,18,1587]
s: 1588
d: [2,1610,1,21,7,0,40,328,330]
d: [2,1611,1,13,4,0,18,1610]
s: 1611
d: [2,2320,2,9,4,2,44,377,1499]
s: 2320
d: [2,1613,1,14,5,0,40,329,315]
s: 1613
d: [2,2338,3,9,4,2,44,1925,1499]
s: 2338
d: [2,1627,1,19,7,0,44,289,276]
d: [2,1654,1,17,6,0,21,1627,279]
a: [3,1657,1,15,5,1654]
s: 1657
d: [2,1631,1,19,7,0,44,289,276]
d: [2,1669,1,17,6,0,21,1631,279]
a: [3,1672,1,15,5,1669]
s: 1672
d: [2,1824,2,3,1,1,40,980,272]
a: [3,1831,2,3,1,1824]
s: 1831
d: [2,1614,1,17,4,0,40,329,320]
s: 1614
d: [2,1825,2,3,1,1,40,980,271]
a: [3,1836,2,3,1,1825]
s: 1836
d: [2,1629,1,22,8,0,44,289,278]
d: [2,1636,1,20,7,0,18,1629]
d: [2,1664,1,18,6,0,21,1636,279]
s: 1664
d: [2,1633,1,22,8,0,44,289,278]
d: [2,1634,1,20,7,0,18,1633]
d: [2,1679,1,18,6,0,21,1634,279]
s: 1679
d: [2,2614,2,13,5,1,40,1679,372]
d: [2,2627,2,10,4,2,21,2614,382]
d: [2,2628,2,8,3,3,21,2627,387]
d: [2,2629,2,6,2,4,21,2628,392]
a: [3,2637,2,3,1,2629]
s: 2637
d: [2,2639,3,0,0,5,52,1188,1106,392,1416,2637,1399]
e: 2639
d: [2,2640,3,0,0,5,52,392,1106,1183,1416,2637,285]
e: 2640
d: [2,2642,3,8,4,1,40,2637,1399]
d: [2,2649,3,6,3,2,21,2642,1416]
d: [2,2650,3,4,2,3,21,2649,392]
d: [2,2651,3,2,1,4,21,2650,1106]
d: [2,2652,3,0,0,5,21,2651,1188]
e: 2652
d: [2,2643,3,8,4,1,40,2637,285]
d: [2,2654,3,6,3,2,21,2643,1416]
d: [2,2655,3,4,2,3,21,2654,1183]
d: [2,2656,3,2,1,4,21,2655,1106]
d: [2,2657,3,0,0,5,21,2656,392]
e: 2657
r: 1179
r: 1188
r: 1192
r: 1201
r: 2637
d: [2,966,1,2,1,1,40,294,691]
d: [2,1445,2,4,2,1,40,1441,307]
d: [2,1447,2,2,1,2,21,1445,685]
d: [2,1454,1,6,3,2,21,1429,387]
d: [2,1428,1,8,4,1,40,284,382]
d: [2,1460,1,6,3,2,21,1428,387]
d: [2,1461,1,4,2,3,21,1460,392]
d: [2,1542,1,10,5,1,40,281,382]
d: [2,2685,1,8,4,2,21,1542,1196]
d: [2,2686,1,6,3,3,21,2685,387]
d: [2,2687,1,4,2,4,21,2686,392]
d: [2,2300,2,7,3,2,44,372,1114]
d: [2,2700,2,5,2,3,21,2300,812]
d: [2,2701,2,2,1,4,21,2700,1951]
f: 2701,2687,2657,2629,2170,1945,1921,1770,1517,1461,1454,1447,1094,1093,1042,1000,967,966,803,693,687,681,352,353,354,261,263,264,252,251,250,247,240,241,242,243,244
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t184_xreal_1
% SZS output start Proof for xreal_1__t184_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(k7_xcmplx_0(X1,X0),1)))) [input t184_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(k7_xcmplx_0(X1,X0),1)))) [negated conjecture 1]
7. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xreal_0(X0)) [input cc1_xreal_0]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
13. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
14. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
19. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
27. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
40. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
44. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc34_xreal_0]
50. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ! [X2] : (v1_xreal_0(X2) => ~(r1_xxreal_0(X1,X2) & ~r1_xxreal_0(k7_xcmplx_0(X2,X0),k7_xcmplx_0(X1,X0)) & r1_xxreal_0(X0,k6_numbers))))) [input l102_xreal_1]
61. ? [X0] : (v1_xreal_0(X0) & v3_xxreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0)) [input rc3_xreal_0]
67. k4_ordinal1 = k5_numbers [input redefinition_k5_numbers]
68. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
80. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
81. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
82. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
90. ! [X0] : (v1_xcmplx_0(X0) => k6_numbers = k7_xcmplx_0(k6_numbers,X0)) [input t5_arithm]
91. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v2_xxreal_0(X1) | v3_xxreal_0(X0) | v1_xboole_0(X1))))) [input t5_real]
93. ! [X0] : (v1_xcmplx_0(X0) => (k6_numbers != X0 => 1 = k7_xcmplx_0(X0,X0))) [input t60_xcmplx_1]
95. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
100. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
112. ? [X0] : (? [X1] : ((~r1_xxreal_0(k7_xcmplx_0(X1,X0),1) & (r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,X1))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
113. ? [X0] : (? [X1] : (~r1_xxreal_0(k7_xcmplx_0(X1,X0),1) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 112]
117. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
118. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
121. ! [X0] : (k6_numbers = k7_xcmplx_0(k6_numbers,X0) | ~v1_xcmplx_0(X0)) [ennf transformation 90]
123. ! [X0] : ((1 = k7_xcmplx_0(X0,X0) | k6_numbers = X0) | ~v1_xcmplx_0(X0)) [ennf transformation 93]
124. ! [X0] : (1 = k7_xcmplx_0(X0,X0) | k6_numbers = X0 | ~v1_xcmplx_0(X0)) [flattening 123]
125. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 13]
127. ! [X0] : (! [X1] : (((v2_xxreal_0(X1) | v3_xxreal_0(X0) | v1_xboole_0(X1)) | ~r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 91]
128. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | v3_xxreal_0(X0) | v1_xboole_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 127]
133. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 82]
134. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 133]
135. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 100]
136. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 135]
143. ! [X0] : (! [X1] : (! [X2] : ((~r1_xxreal_0(X1,X2) | r1_xxreal_0(k7_xcmplx_0(X2,X0),k7_xcmplx_0(X1,X0)) | ~r1_xxreal_0(X0,k6_numbers)) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 50]
144. ! [X0] : (! [X1] : (! [X2] : (~r1_xxreal_0(X1,X2) | r1_xxreal_0(k7_xcmplx_0(X2,X0),k7_xcmplx_0(X1,X0)) | ~r1_xxreal_0(X0,k6_numbers) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 143]
148. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 95]
149. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
151. ! [X0] : (v1_xreal_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 7]
160. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 19]
161. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 160]
162. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 14]
163. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 162]
182. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 44]
183. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 182]
198. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 27]
199. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 198]
200. ? [X0] : (? [X1] : (~r1_xxreal_0(k7_xcmplx_0(X1,X0),1) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~r1_xxreal_0(k7_xcmplx_0(X1,sK0),1) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
201. ? [X1] : (~r1_xxreal_0(k7_xcmplx_0(X1,sK0),1) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) => (~r1_xxreal_0(k7_xcmplx_0(sK1,sK0),1) & r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) [choice axiom]
202. (~r1_xxreal_0(k7_xcmplx_0(sK1,sK0),1) & r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 113,201,200]
228. ? [X0] : (v1_xreal_0(X0) & v3_xxreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0)) => (v1_xreal_0(sK14) & v3_xxreal_0(sK14) & v1_xxreal_0(sK14) & v1_xcmplx_0(sK14)) [choice axiom]
229. v1_xreal_0(sK14) & v3_xxreal_0(sK14) & v1_xxreal_0(sK14) & v1_xcmplx_0(sK14) [skolemisation 61,228]
240. v1_xreal_0(sK0) [cnf transformation 202] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
241. v1_xreal_0(sK1) [cnf transformation 202] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
242. r1_xxreal_0(sK0,sK1) [cnf transformation 202] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
243. r1_xxreal_0(sK1,k6_numbers) [cnf transformation 202] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
244. ~r1_xxreal_0(k7_xcmplx_0(sK1,sK0),1) [cnf transformation 202] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
247. v1_xboole_0(k1_xboole_0) [cnf transformation 40] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
250. k4_ordinal1 = k5_numbers [cnf transformation 67] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
251. k6_numbers = k1_xboole_0 [cnf transformation 68] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
252. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 81] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
261. v2_xxreal_0(1) [cnf transformation 80] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
263. m1_subset_1(1,k5_numbers) [cnf transformation 80] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
264. m1_subset_1(1,k1_numbers) [cnf transformation 80] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
271. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 117] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
272. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 118] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
276. k6_numbers = k7_xcmplx_0(k6_numbers,X0) | ~v1_xcmplx_0(X0) [cnf transformation 121] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
278. 1 = k7_xcmplx_0(X0,X0) | k6_numbers = X0 | ~v1_xcmplx_0(X0) [cnf transformation 124] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
279. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 125] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
281. ~r1_xxreal_0(X0,X1) | v3_xxreal_0(X0) | v1_xboole_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 128] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
284. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 134] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
285. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 136] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
289. r1_xxreal_0(k7_xcmplx_0(X2,X0),k7_xcmplx_0(X1,X0)) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,k6_numbers) | ~v1_xreal_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 144] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
293. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 148] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
294. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 149] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
296. ~m1_subset_1(X0,k1_numbers) | v1_xreal_0(X0) [cnf transformation 151] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
306. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 161] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
307. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 163] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
326. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 183] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
335. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 199] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
352. v1_xxreal_0(sK14) [cnf transformation 229] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
353. v3_xxreal_0(sK14) [cnf transformation 229] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
354. v1_xreal_0(sK14) [cnf transformation 229] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
370. 1 <=> r1_xxreal_0(k7_xcmplx_0(sK1,sK0),1) [avatar definition]
372. ~r1_xxreal_0(k7_xcmplx_0(sK1,sK0),1) <- (~1) [avatar component clause 370] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
373. ~1 [avatar split clause 244,370]
375. 2 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
377. r1_xxreal_0(sK1,k6_numbers) <- (2) [avatar component clause 375] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
378. 2 [avatar split clause 243,375]
380. 3 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
382. r1_xxreal_0(sK0,sK1) <- (3) [avatar component clause 380] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
383. 3 [avatar split clause 242,380]
385. 4 <=> v1_xreal_0(sK1) [avatar definition]
387. v1_xreal_0(sK1) <- (4) [avatar component clause 385] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
388. 4 [avatar split clause 241,385]
390. 5 <=> v1_xreal_0(sK0) [avatar definition]
392. v1_xreal_0(sK0) <- (5) [avatar component clause 390] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
393. 5 [avatar split clause 240,390]
405. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
407. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 405] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
408. 8 [avatar split clause 247,405]
420. 11 <=> k4_ordinal1 = k5_numbers [avatar definition]
422. k4_ordinal1 = k5_numbers <- (11) [avatar component clause 420] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
423. 11 [avatar split clause 250,420]
425. 12 <=> k6_numbers = k1_xboole_0 [avatar definition]
427. k6_numbers = k1_xboole_0 <- (12) [avatar component clause 425] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
428. 12 [avatar split clause 251,425]
430. 13 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
432. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (13) [avatar component clause 430] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
433. 13 [avatar split clause 252,430]
475. 22 <=> m1_subset_1(1,k1_numbers) [avatar definition]
477. m1_subset_1(1,k1_numbers) <- (22) [avatar component clause 475] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
478. 22 [avatar split clause 264,475]
480. 23 <=> m1_subset_1(1,k5_numbers) [avatar definition]
482. m1_subset_1(1,k5_numbers) <- (23) [avatar component clause 480] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
483. 23 [avatar split clause 263,480]
490. 25 <=> v2_xxreal_0(1) [avatar definition]
492. v2_xxreal_0(1) <- (25) [avatar component clause 490] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
493. 25 [avatar split clause 261,490]
592. 46 <=> v1_xreal_0(sK14) [avatar definition]
594. v1_xreal_0(sK14) <- (46) [avatar component clause 592] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
595. 46 [avatar split clause 354,592]
597. 47 <=> v3_xxreal_0(sK14) [avatar definition]
599. v3_xxreal_0(sK14) <- (47) [avatar component clause 597] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
600. 47 [avatar split clause 353,597]
602. 48 <=> v1_xxreal_0(sK14) [avatar definition]
604. v1_xxreal_0(sK14) <- (48) [avatar component clause 602] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
605. 48 [avatar split clause 352,602]
681. v1_xboole_0(k6_numbers) <- (8, 12) [backward demodulation 407,427] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
683. 64 <=> v1_xboole_0(k6_numbers) [avatar definition]
685. v1_xboole_0(k6_numbers) <- (64) [avatar component clause 683] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
686. 64 | ~8 | ~12 [avatar split clause 681,425,405,683]
687. m1_subset_1(k6_numbers,k4_ordinal1) <- (12, 13) [forward demodulation 432,427] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
689. 65 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
691. m1_subset_1(k6_numbers,k4_ordinal1) <- (65) [avatar component clause 689] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
692. 65 | ~12 | ~13 [avatar split clause 687,430,425,689]
693. m1_subset_1(1,k4_ordinal1) <- (11, 23) [forward demodulation 482,422] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
695. 66 <=> m1_subset_1(1,k4_ordinal1) [avatar definition]
697. m1_subset_1(1,k4_ordinal1) <- (66) [avatar component clause 695] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
698. 66 | ~11 | ~23 [avatar split clause 693,480,420,695]
803. v1_xcmplx_0(sK0) <- (5) [resolution 279,392] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
810. 84 <=> v1_xcmplx_0(sK0) [avatar definition]
812. v1_xcmplx_0(sK0) <- (84) [avatar component clause 810] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
813. 84 | ~5 [avatar split clause 803,390,810]
881. 95 <=> v7_ordinal1(k6_numbers) [avatar definition]
883. v7_ordinal1(k6_numbers) <- (95) [avatar component clause 881] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
965. k6_numbers = X0 | ~v1_xboole_0(X0) <- (12) [forward demodulation 293,427] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
966. v7_ordinal1(k6_numbers) <- (65) [resolution 294,691] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
967. v7_ordinal1(1) <- (66) [resolution 294,697] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
970. 95 | ~65 [avatar split clause 966,689,881]
972. 109 <=> v7_ordinal1(1) [avatar definition]
974. v7_ordinal1(1) <- (109) [avatar component clause 972] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
975. 109 | ~66 [avatar split clause 967,695,972]
1000. v1_xreal_0(1) <- (22) [resolution 296,477] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1004. 114 <=> v1_xreal_0(1) [avatar definition]
1006. v1_xreal_0(1) <- (114) [avatar component clause 1004] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1007. 114 | ~22 [avatar split clause 1000,475,1004]
1040. ~v2_xxreal_0(sK14) | ~v1_xxreal_0(sK14) <- (47) [resolution 306,599] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:3,thDist:-3}
1042. ~v2_xxreal_0(sK14) <- (47, 48) [subsumption resolution 1040,604] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:3,thDist:-3}
1044. 119 <=> v2_xxreal_0(sK14) [avatar definition]
1046. ~v2_xxreal_0(sK14) <- (~119) [avatar component clause 1044] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
1047. ~119 | ~47 | ~48 [avatar split clause 1042,602,597,1044]
1093. v1_xxreal_0(k6_numbers) <- (95) [resolution 883,272] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1094. v1_xreal_0(k6_numbers) <- (95) [resolution 883,271] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1099. 125 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1101. v1_xxreal_0(k6_numbers) <- (125) [avatar component clause 1099] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
1102. 125 | ~95 [avatar split clause 1093,881,1099]
1104. 126 <=> v1_xreal_0(k6_numbers) [avatar definition]
1106. v1_xreal_0(k6_numbers) <- (126) [avatar component clause 1104] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
1107. 126 | ~95 [avatar split clause 1094,881,1104]
1114. k7_xcmplx_0(X0,X1) = X0 | ~v1_xcmplx_0(X1) | ~v1_xboole_0(X0) <- (12) [superposition 276,965] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1182. 135 <=> v2_xxreal_0(sK0) [avatar definition]
1183. ~v2_xxreal_0(sK0) <- (~135) [avatar component clause 1182] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1186. 136 <=> v3_xxreal_0(sK0) [avatar definition]
1191. 137 <=> v1_xboole_0(sK1) [avatar definition]
1195. 138 <=> v2_xxreal_0(sK1) [avatar definition]
1196. ~v2_xxreal_0(sK1) <- (~138) [avatar component clause 1195] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1375. 166 <=> v1_xxreal_0(1) [avatar definition]
1376. v1_xxreal_0(1) <- (166) [avatar component clause 1375] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1414. 171 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1416. ~v3_xxreal_0(k6_numbers) <- (~171) [avatar component clause 1414] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1428. ~v2_xxreal_0(sK0) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (3) [resolution 284,382] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1429. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (2) [resolution 284,377] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1439. 173 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1440. ~v2_xxreal_0(k6_numbers) <- (~173) [avatar component clause 1439] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1441. v2_xxreal_0(k6_numbers) <- (173) [avatar component clause 1439] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1445. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (173) [resolution 1441,307] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1447. ~v1_xxreal_0(k6_numbers) <- (64, 173) [subsumption resolution 1445,685] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1448. ~125 | ~64 | ~173 [avatar split clause 1447,1439,683,1099]
1454. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 4) [subsumption resolution 1429,387] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1455. ~126 | 173 | ~138 | ~2 | ~4 [avatar split clause 1454,385,375,1195,1439,1104]
1460. ~v2_xxreal_0(sK0) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) <- (3, 4) [subsumption resolution 1428,387] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1461. ~v2_xxreal_0(sK0) | v2_xxreal_0(sK1) <- (3, 4, 5) [subsumption resolution 1460,392] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1462. 138 | ~135 | ~3 | ~4 | ~5 [avatar split clause 1461,390,385,380,1182,1195]
1511. ~v3_xxreal_0(k6_numbers) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) | ~v1_xreal_0(k6_numbers) | v2_xxreal_0(k6_numbers) | ~v1_xcmplx_0(X0) [superposition 326,276] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
1516. ~v3_xxreal_0(k6_numbers) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) | ~v1_xreal_0(k6_numbers) | ~v1_xcmplx_0(X0) <- (~173) [subsumption resolution 1511,1440] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1517. ~v3_xxreal_0(k6_numbers) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) | ~v1_xreal_0(k6_numbers) <- (~173) [subsumption resolution 1516,279] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1519. 178 <=> ! [X0] : (~v1_xreal_0(X0) | v2_xxreal_0(X0)) [avatar definition]
1520. ~v1_xreal_0(X0) | v2_xxreal_0(X0) <- (178) [avatar component clause 1519] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1521. ~126 | 178 | ~171 | 173 [avatar split clause 1517,1439,1414,1519,1104]
1542. v3_xxreal_0(sK0) | v1_xboole_0(sK1) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (3) [resolution 281,382] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1633. r1_xxreal_0(k7_xcmplx_0(X6,X5),1) | ~r1_xxreal_0(X5,X6) | ~r1_xxreal_0(X5,k6_numbers) | ~v1_xreal_0(X6) | ~v1_xreal_0(X5) | ~v1_xreal_0(X5) | k6_numbers = X5 | ~v1_xcmplx_0(X5) [superposition 289,278] {a:1,w:22,wCS:8800,sine:1,thAx:0,allAx:2,thDist:-2}
1634. r1_xxreal_0(k7_xcmplx_0(X6,X5),1) | ~r1_xxreal_0(X5,X6) | ~r1_xxreal_0(X5,k6_numbers) | ~v1_xreal_0(X6) | ~v1_xreal_0(X5) | k6_numbers = X5 | ~v1_xcmplx_0(X5) [duplicate literal removal 1633] {a:1,w:20,wCS:8000,sine:1,thAx:0,allAx:2,thDist:-2}
1647. 186 <=> r1_xxreal_0(1,k6_numbers) [avatar definition]
1648. r1_xxreal_0(1,k6_numbers) <- (186) [avatar component clause 1647] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1649. ~r1_xxreal_0(1,k6_numbers) <- (~186) [avatar component clause 1647] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1679. r1_xxreal_0(k7_xcmplx_0(X6,X5),1) | ~r1_xxreal_0(X5,X6) | ~r1_xxreal_0(X5,k6_numbers) | ~v1_xreal_0(X6) | ~v1_xreal_0(X5) | k6_numbers = X5 [subsumption resolution 1634,279] {a:1,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1770. v1_xxreal_0(1) <- (109) [resolution 974,272] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1777. 166 | ~109 [avatar split clause 1770,972,1375]
1918. r1_xxreal_0(k6_numbers,1) | ~v1_xxreal_0(1) | ~v1_xxreal_0(k6_numbers) <- (~186) [resolution 1649,335] {a:2,w:7,wCS:2800,sine:1,thAx:0,allAx:3,thDist:-3}
1920. r1_xxreal_0(k6_numbers,1) | ~v1_xxreal_0(k6_numbers) <- (166, ~186) [subsumption resolution 1918,1376] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1921. r1_xxreal_0(k6_numbers,1) <- (125, 166, ~186) [subsumption resolution 1920,1101] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1923. 221 <=> r1_xxreal_0(k6_numbers,1) [avatar definition]
1925. r1_xxreal_0(k6_numbers,1) <- (221) [avatar component clause 1923] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1926. 221 | ~125 | ~166 | 186 [avatar split clause 1921,1647,1375,1099,1923]
1931. ~v2_xxreal_0(1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(1) <- (186) [resolution 1648,284] {a:2,w:8,wCS:3200,sine:1,thAx:0,allAx:3,thDist:-3}
1942. v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(1) <- (25, 186) [subsumption resolution 1931,492] {a:2,w:6,wCS:2400,sine:1,thAx:0,allAx:3,thDist:-3}
1943. ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(1) <- (25, ~173, 186) [subsumption resolution 1942,1440] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1944. ~v1_xreal_0(1) <- (25, 126, ~173, 186) [subsumption resolution 1943,1106] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1945. $false <- (25, 114, 126, ~173, 186) [subsumption resolution 1944,1006] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1946. ~25 | ~114 | ~126 | 173 | ~186 [avatar contradiction clause 1945]
1951. r1_xxreal_0(X0,1) | ~v1_xboole_0(X0) <- (12, 221) [superposition 1925,965] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2156. v2_xxreal_0(sK14) <- (46, 178) [resolution 1520,594] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2170. $false <- (46, ~119, 178) [subsumption resolution 2156,1046] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2171. ~46 | 119 | ~178 [avatar contradiction clause 2170]
2300. ~r1_xxreal_0(sK1,1) | ~v1_xcmplx_0(sK0) | ~v1_xboole_0(sK1) <- (~1, 12) [superposition 372,1114] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2614. ~r1_xxreal_0(sK0,sK1) | ~r1_xxreal_0(sK0,k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) | k6_numbers = sK0 <- (~1) [resolution 1679,372] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2627. ~r1_xxreal_0(sK0,k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) | k6_numbers = sK0 <- (~1, 3) [subsumption resolution 2614,382] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2628. ~r1_xxreal_0(sK0,k6_numbers) | ~v1_xreal_0(sK0) | k6_numbers = sK0 <- (~1, 3, 4) [subsumption resolution 2627,387] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2629. ~r1_xxreal_0(sK0,k6_numbers) | k6_numbers = sK0 <- (~1, 3, 4, 5) [subsumption resolution 2628,392] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2631. 248 <=> k6_numbers = sK0 [avatar definition]
2635. 249 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
2637. ~r1_xxreal_0(sK0,k6_numbers) <- (~249) [avatar component clause 2635] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2638. 248 | ~249 | 1 | ~3 | ~4 | ~5 [avatar split clause 2629,390,385,380,370,2635,2631]
2643. v3_xxreal_0(k6_numbers) | v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (~249) [resolution 2637,285] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2654. v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (~171, ~249) [subsumption resolution 2643,1416] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2655. ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (~135, ~171, ~249) [subsumption resolution 2654,1183] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2656. ~v1_xreal_0(sK0) <- (126, ~135, ~171, ~249) [subsumption resolution 2655,1106] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2657. $false <- (5, 126, ~135, ~171, ~249) [subsumption resolution 2656,392] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2658. ~5 | ~126 | 135 | 171 | 249 [avatar contradiction clause 2657]
2662. k6_numbers != sK0 | v3_xxreal_0(k6_numbers) | ~v3_xxreal_0(sK0) [theory tautology sat conflict] {a:0,w:7,wCS:2800,sine:255,thAx:0,allAx:1,thDist:-1}
2685. v3_xxreal_0(sK0) | v1_xboole_0(sK1) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (3, ~138) [subsumption resolution 1542,1196] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2686. v3_xxreal_0(sK0) | v1_xboole_0(sK1) | ~v1_xreal_0(sK0) <- (3, 4, ~138) [subsumption resolution 2685,387] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2687. v3_xxreal_0(sK0) | v1_xboole_0(sK1) <- (3, 4, 5, ~138) [subsumption resolution 2686,392] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2688. 137 | 136 | ~3 | ~4 | ~5 | 138 [avatar split clause 2687,1195,390,385,380,1186,1191]
2700. ~r1_xxreal_0(sK1,1) | ~v1_xboole_0(sK1) <- (~1, 12, 84) [subsumption resolution 2300,812] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2701. ~v1_xboole_0(sK1) <- (~1, 12, 84, 221) [subsumption resolution 2700,1951] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2702. ~137 | 1 | ~12 | ~84 | ~221 [avatar split clause 2701,1923,810,425,370,1191]
2732. $false [avatar sat refutation 373,378,383,388,393,408,423,428,433,478,483,493,595,600,605,686,692,698,813,970,975,1007,1047,1102,1107,1448,1455,1462,1521,1777,1926,1946,2171,2638,2658,2662,2688,2702] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1792,thDist:-1792}
% SZS output end Proof for xreal_1__t184_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 99
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 129
% Generated clauses: 1480
% Active clauses: 291
% Passive clauses: 443
% Final active clauses: 255
% Final passive clauses: 106
% Main loop iterations started: 312
% 
% >>> Simplifying Inferences
% Duplicate literals: 153
% Fw subsumption resolutions: 236
% Fw demodulations: 8
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 76
% Forward subsumptions: 303
% Backward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 539
% Unit resulting resolution: 11
% Factoring: 1
% Forward superposition: 121
% Backward superposition: 46
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 66
% Split components: 158
% Sat splitting refutations: 23
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 310
% SAT solver unit clauses: 64
% SAT solver binary clauses: 126
% 
% Memory used [KB]: 6652
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% congruence closure: 0.002 s
% backward demodulation index maintenance: 0.001 s
% forward subsumption: 0.002 s
% unit resulting resolution: 0.001 s
% superposition: 0.003 s
% other: 0.016 s ( own 0.005 s ) 

