i: [1,38,0,2,1,1,0,0] t5_orders_2
a: [3,88,0,2,1,38]
s: 88
i: [1,37,0,2,1,1,0,0] t5_orders_2
a: [3,93,0,2,1,37]
s: 93
i: [1,36,0,2,1,1,0,0] t5_orders_2
a: [3,98,0,2,1,36]
s: 98
i: [1,44,0,4,1,1,0,0] t5_orders_2
a: [3,58,0,4,1,44]
s: 58
i: [1,43,0,4,1,1,0,0] t5_orders_2
a: [3,63,0,4,1,43]
s: 63
i: [1,42,0,4,1,1,0,0] t5_orders_2
a: [3,68,0,4,1,42]
s: 68
i: [1,41,0,4,1,1,0,0] t5_orders_2
a: [3,73,0,4,1,41]
s: 73
i: [1,40,0,4,1,1,0,0] t5_orders_2
a: [3,78,0,4,1,40]
s: 78
i: [1,39,0,4,1,1,0,0] t5_orders_2
a: [3,83,0,4,1,39]
s: 83
i: [1,52,0,2,1,0,0,1] existence_l1_orders_2
a: [3,103,0,2,1,52]
s: 103
i: [1,50,0,4,1,0,0,1] existence_m1_subset_1
s: 50
i: [1,53,0,14,4,0,0,255] d6_orders_2
d: [2,54,0,10,3,0,18,53]
s: 54
i: [1,51,0,14,4,0,0,1] irreflexivity_r2_orders_2
s: 51
i: [1,45,0,18,5,0,0,255] d6_orders_2
s: 45
i: [1,48,0,20,6,0,0,1] t4_orders_2
s: 48
i: [1,47,0,21,6,0,0,255] d6_orders_2
s: 47
d: [2,105,1,14,4,1,40,45,63]
d: [2,107,1,10,3,2,21,105,73]
d: [2,108,1,6,2,3,21,107,78]
d: [2,109,1,4,1,4,21,108,88]
a: [3,113,1,4,1,109]
s: 113
i: [1,49,0,28,8,0,0,1] t3_orders_2
s: 49
d: [2,106,1,14,4,1,40,45,68]
d: [2,115,1,10,3,2,21,106,78]
d: [2,116,1,6,2,3,21,115,83]
d: [2,117,1,4,1,4,21,116,88]
a: [3,121,1,4,1,117]
s: 121
d: [2,123,1,16,5,1,40,48,63]
d: [2,125,1,12,4,2,21,123,78]
d: [2,126,1,8,3,3,21,125,73]
d: [2,127,1,6,2,4,21,126,88]
d: [2,128,1,4,1,5,21,127,93]
a: [3,132,1,4,1,128]
s: 132
d: [2,124,1,16,5,1,40,48,68]
d: [2,134,1,12,4,2,21,124,83]
d: [2,135,1,8,3,3,21,134,78]
d: [2,136,1,6,2,4,21,135,88]
d: [2,137,1,4,1,5,21,136,93]
a: [3,141,1,4,1,137]
s: 141
d: [2,144,2,24,7,1,40,49,113]
d: [2,145,2,20,6,2,21,144,73]
d: [2,146,2,16,5,3,21,145,78]
d: [2,147,2,14,4,4,21,146,88]
d: [2,148,2,12,3,5,21,147,98]
s: 148
d: [2,149,2,24,7,1,40,121,49]
d: [2,151,2,20,6,2,21,149,78]
d: [2,152,2,16,5,3,21,151,83]
d: [2,153,2,14,4,4,21,152,88]
d: [2,154,2,12,3,5,21,153,98]
s: 154
d: [2,155,3,8,2,6,40,148,121]
d: [2,156,3,4,1,7,21,155,83]
a: [3,160,3,4,1,156]
s: 160
d: [2,163,4,17,5,1,40,160,47]
d: [2,168,4,13,4,2,21,163,58]
d: [2,169,4,9,3,3,21,168,73]
d: [2,170,4,5,2,4,21,169,83]
d: [2,171,4,3,1,5,21,170,88]
f: 171,156,128,117,109,36,37,38,39,40,41,42,43,44
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_2__t5_orders_2
% SZS output start Proof for orders_2__t5_orders_2
1. ! [X0] : ((l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ((r2_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2)) => r2_orders_2(X0,X1,X3)))))) [input t5_orders_2]
2. ~! [X0] : ((l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ((r2_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2)) => r2_orders_2(X0,X1,X3)))))) [negated conjecture 1]
3. ! [X0] : (l1_orders_2(X0) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_orders_2(X0,X1,X2) <=> (X1 != X2 & r1_orders_2(X0,X1,X2)))))) [input d6_orders_2]
12. ! [X0] : ((l1_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ((r1_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X2)) => r1_orders_2(X0,X1,X3)))))) [input t3_orders_2]
13. ! [X0] : ((l1_orders_2(X0) & v5_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ~(r2_orders_2(X0,X2,X1) & r2_orders_2(X0,X1,X2))))) [input t4_orders_2]
16. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r2_orders_2(X0,X1,X3) & (r2_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0))) [ennf transformation 2]
17. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_orders_2(X0,X1,X3) & r2_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0)) [flattening 16]
18. ! [X0] : (! [X1] : (! [X2] : ((r2_orders_2(X0,X1,X2) <=> (X1 != X2 & r1_orders_2(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0)) [ennf transformation 3]
19. ! [X0] : (! [X1] : (! [X2] : ((~r2_orders_2(X0,X2,X1) | ~r2_orders_2(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v5_orders_2(X0))) [ennf transformation 13]
20. ! [X0] : (! [X1] : (! [X2] : (~r2_orders_2(X0,X2,X1) | ~r2_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v5_orders_2(X0)) [flattening 19]
21. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_orders_2(X0,X1,X3) | (~r1_orders_2(X0,X2,X3) | ~r1_orders_2(X0,X1,X2))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v4_orders_2(X0))) [ennf transformation 12]
22. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (r1_orders_2(X0,X1,X3) | ~r1_orders_2(X0,X2,X3) | ~r1_orders_2(X0,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v4_orders_2(X0)) [flattening 21]
25. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_orders_2(X0,X1,X3) & r2_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r2_orders_2(sK0,X1,X3) & r2_orders_2(sK0,X2,X3) & r2_orders_2(sK0,X1,X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v5_orders_2(sK0) & v4_orders_2(sK0)) [choice axiom]
26. ? [X1] : (? [X2] : (? [X3] : (~r2_orders_2(sK0,X1,X3) & r2_orders_2(sK0,X2,X3) & r2_orders_2(sK0,X1,X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (~r2_orders_2(sK0,sK1,X3) & r2_orders_2(sK0,X2,X3) & r2_orders_2(sK0,sK1,X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
27. ? [X2] : (? [X3] : (~r2_orders_2(sK0,sK1,X3) & r2_orders_2(sK0,X2,X3) & r2_orders_2(sK0,sK1,X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (~r2_orders_2(sK0,sK1,X3) & r2_orders_2(sK0,sK2,X3) & r2_orders_2(sK0,sK1,sK2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
28. ? [X3] : (~r2_orders_2(sK0,sK1,X3) & r2_orders_2(sK0,sK2,X3) & r2_orders_2(sK0,sK1,sK2) & m1_subset_1(X3,u1_struct_0(sK0))) => (~r2_orders_2(sK0,sK1,sK3) & r2_orders_2(sK0,sK2,sK3) & r2_orders_2(sK0,sK1,sK2) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
29. (((~r2_orders_2(sK0,sK1,sK3) & r2_orders_2(sK0,sK2,sK3) & r2_orders_2(sK0,sK1,sK2) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v5_orders_2(sK0) & v4_orders_2(sK0) [skolemisation 17,28,27,26,25]
30. ! [X0] : (! [X1] : (! [X2] : (((r2_orders_2(X0,X1,X2) | (X1 = X2 | ~r1_orders_2(X0,X1,X2))) & ((X1 != X2 & r1_orders_2(X0,X1,X2)) | ~r2_orders_2(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0)) [nnf transformation 18]
31. ! [X0] : (! [X1] : (! [X2] : (((r2_orders_2(X0,X1,X2) | X1 = X2 | ~r1_orders_2(X0,X1,X2)) & ((X1 != X2 & r1_orders_2(X0,X1,X2)) | ~r2_orders_2(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0)) [flattening 30]
36. v4_orders_2(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
37. v5_orders_2(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
38. l1_orders_2(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
39. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
40. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
41. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
42. r2_orders_2(sK0,sK1,sK2) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
43. r2_orders_2(sK0,sK2,sK3) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
44. ~r2_orders_2(sK0,sK1,sK3) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
45. ~r2_orders_2(X0,X1,X2) | r1_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 31] {a:0,w:18,wCS:7200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
47. ~r1_orders_2(X0,X1,X2) | X1 = X2 | r2_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 31] {a:0,w:21,wCS:8400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
48. ~r2_orders_2(X0,X2,X1) | ~r2_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v5_orders_2(X0) [cnf transformation 20] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. ~r1_orders_2(X0,X2,X3) | r1_orders_2(X0,X1,X3) | ~r1_orders_2(X0,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v4_orders_2(X0) [cnf transformation 22] {a:0,w:28,wCS:11200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. 1 <=> r2_orders_2(sK0,sK1,sK3) [avatar definition]
58. ~r2_orders_2(sK0,sK1,sK3) <- (~1) [avatar component clause 56] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
59. ~1 [avatar split clause 44,56]
61. 2 <=> r2_orders_2(sK0,sK2,sK3) [avatar definition]
63. r2_orders_2(sK0,sK2,sK3) <- (2) [avatar component clause 61] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
64. 2 [avatar split clause 43,61]
66. 3 <=> r2_orders_2(sK0,sK1,sK2) [avatar definition]
68. r2_orders_2(sK0,sK1,sK2) <- (3) [avatar component clause 66] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
69. 3 [avatar split clause 42,66]
71. 4 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
73. m1_subset_1(sK3,u1_struct_0(sK0)) <- (4) [avatar component clause 71] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
74. 4 [avatar split clause 41,71]
76. 5 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
78. m1_subset_1(sK2,u1_struct_0(sK0)) <- (5) [avatar component clause 76] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
79. 5 [avatar split clause 40,76]
81. 6 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
83. m1_subset_1(sK1,u1_struct_0(sK0)) <- (6) [avatar component clause 81] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
84. 6 [avatar split clause 39,81]
86. 7 <=> l1_orders_2(sK0) [avatar definition]
88. l1_orders_2(sK0) <- (7) [avatar component clause 86] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
89. 7 [avatar split clause 38,86]
91. 8 <=> v5_orders_2(sK0) [avatar definition]
93. v5_orders_2(sK0) <- (8) [avatar component clause 91] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
94. 8 [avatar split clause 37,91]
96. 9 <=> v4_orders_2(sK0) [avatar definition]
98. v4_orders_2(sK0) <- (9) [avatar component clause 96] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
99. 9 [avatar split clause 36,96]
105. r1_orders_2(sK0,sK2,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (2) [resolution 45,63] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
106. r1_orders_2(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (3) [resolution 45,68] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
107. r1_orders_2(sK0,sK2,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (2, 4) [subsumption resolution 105,73] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
108. r1_orders_2(sK0,sK2,sK3) | ~l1_orders_2(sK0) <- (2, 4, 5) [subsumption resolution 107,78] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
109. r1_orders_2(sK0,sK2,sK3) <- (2, 4, 5, 7) [subsumption resolution 108,88] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
111. 11 <=> r1_orders_2(sK0,sK2,sK3) [avatar definition]
113. r1_orders_2(sK0,sK2,sK3) <- (11) [avatar component clause 111] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
114. 11 | ~2 | ~4 | ~5 | ~7 [avatar split clause 109,86,76,71,61,111]
115. r1_orders_2(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (3, 5) [subsumption resolution 106,78] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
116. r1_orders_2(sK0,sK1,sK2) | ~l1_orders_2(sK0) <- (3, 5, 6) [subsumption resolution 115,83] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
117. r1_orders_2(sK0,sK1,sK2) <- (3, 5, 6, 7) [subsumption resolution 116,88] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
119. 12 <=> r1_orders_2(sK0,sK1,sK2) [avatar definition]
121. r1_orders_2(sK0,sK1,sK2) <- (12) [avatar component clause 119] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
122. 12 | ~3 | ~5 | ~6 | ~7 [avatar split clause 117,86,81,76,66,119]
123. ~r2_orders_2(sK0,sK3,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v5_orders_2(sK0) <- (2) [resolution 48,63] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
125. ~r2_orders_2(sK0,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v5_orders_2(sK0) <- (2, 5) [subsumption resolution 123,78] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
126. ~r2_orders_2(sK0,sK3,sK2) | ~l1_orders_2(sK0) | ~v5_orders_2(sK0) <- (2, 4, 5) [subsumption resolution 125,73] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
127. ~r2_orders_2(sK0,sK3,sK2) | ~v5_orders_2(sK0) <- (2, 4, 5, 7) [subsumption resolution 126,88] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
128. ~r2_orders_2(sK0,sK3,sK2) <- (2, 4, 5, 7, 8) [subsumption resolution 127,93] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
130. 13 <=> r2_orders_2(sK0,sK3,sK2) [avatar definition]
133. ~13 | ~2 | ~4 | ~5 | ~7 | ~8 [avatar split clause 128,91,86,76,71,61,130]
144. r1_orders_2(sK0,X0,sK3) | ~r1_orders_2(sK0,X0,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v4_orders_2(sK0) <- (11) [resolution 49,113] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
145. r1_orders_2(sK0,X0,sK3) | ~r1_orders_2(sK0,X0,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v4_orders_2(sK0) <- (4, 11) [subsumption resolution 144,73] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
146. r1_orders_2(sK0,X0,sK3) | ~r1_orders_2(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v4_orders_2(sK0) <- (4, 5, 11) [subsumption resolution 145,78] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
147. r1_orders_2(sK0,X0,sK3) | ~r1_orders_2(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v4_orders_2(sK0) <- (4, 5, 7, 11) [subsumption resolution 146,88] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
148. ~r1_orders_2(sK0,X0,sK2) | r1_orders_2(sK0,X0,sK3) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (4, 5, 7, 9, 11) [subsumption resolution 147,98] {a:2,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
155. r1_orders_2(sK0,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) <- (4, 5, 7, 9, 11, 12) [resolution 148,121] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
156. r1_orders_2(sK0,sK1,sK3) <- (4, 5, 6, 7, 9, 11, 12) [subsumption resolution 155,83] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
158. 15 <=> r1_orders_2(sK0,sK1,sK3) [avatar definition]
160. r1_orders_2(sK0,sK1,sK3) <- (15) [avatar component clause 158] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
161. 15 | ~4 | ~5 | ~6 | ~7 | ~9 | ~11 | ~12 [avatar split clause 156,119,111,96,86,81,76,71,158]
163. sK1 = sK3 | r2_orders_2(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (15) [resolution 160,47] {a:4,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
168. sK1 = sK3 | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (~1, 15) [subsumption resolution 163,58] {a:4,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
169. sK1 = sK3 | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (~1, 4, 15) [subsumption resolution 168,73] {a:4,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
170. sK1 = sK3 | ~l1_orders_2(sK0) <- (~1, 4, 6, 15) [subsumption resolution 169,83] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
171. sK1 = sK3 <- (~1, 4, 6, 7, 15) [subsumption resolution 170,88] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
173. 16 <=> sK1 = sK3 [avatar definition]
176. 16 | 1 | ~4 | ~6 | ~7 | ~15 [avatar split clause 171,158,86,81,71,56,173]
177. sK1 != sK3 | r2_orders_2(sK0,sK3,sK2) | ~r2_orders_2(sK0,sK1,sK2) [theory tautology sat conflict] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
178. $false [avatar sat refutation 59,64,69,74,79,84,89,94,99,114,122,133,161,176,177] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:145,thDist:-145}
% SZS output end Proof for orders_2__t5_orders_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 75
% Active clauses: 23
% Passive clauses: 25
% Final active clauses: 23
% Final passive clauses: 1
% Main loop iterations started: 24
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 31
% 
% >>> Deletion Inferences
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 11
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 19
% SAT solver unit clauses: 10
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

