i: [1,36,0,2,1,1,0,0] t99_relat_1
a: [3,64,0,2,1,36]
s: 64
i: [1,37,0,3,1,1,0,0] t99_relat_1
a: [3,59,0,3,1,37]
s: 59
i: [1,38,0,9,1,1,0,0] t99_relat_1
a: [3,54,0,9,1,38]
s: 54
i: [1,45,0,3,1,0,0,1] reflexivity_r1_tarski
s: 45
i: [1,44,0,4,1,0,0,255] existence_m1_subset_1
s: 44
i: [1,46,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 46
i: [1,48,0,6,2,0,0,2] fc1_relat_1
s: 48
i: [1,49,0,6,2,0,0,1] dt_k6_relat_1
s: 49
i: [1,41,0,7,2,0,0,255] t3_subset
s: 41
i: [1,42,0,7,2,0,0,255] t3_subset
s: 42
i: [1,47,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 47
d: [2,67,1,5,1,0,40,41,44]
s: 67
i: [1,39,0,8,3,0,0,255] cc2_relat_1
s: 39
d: [2,69,1,6,2,0,44,48,47]
s: 69
i: [1,40,0,8,2,0,0,1] t28_xboole_1
s: 40
i: [1,50,0,11,2,0,0,1] rd7_relat_1
s: 50
d: [2,70,1,6,2,0,44,48,47]
s: 70
i: [1,43,0,13,2,0,0,1] t96_relat_1
s: 43
d: [2,71,1,6,2,0,40,39,44]
s: 71
d: [2,72,1,7,3,0,40,39,42]
s: 72
d: [2,76,1,8,2,0,44,40,47]
s: 76
d: [2,96,2,4,2,1,40,72,59]
a: [3,106,2,2,1,96]
s: 106
d: [2,77,1,8,2,0,44,40,47]
s: 77
d: [2,78,2,7,3,0,44,69,40]
s: 78
d: [2,79,1,8,2,0,44,47,40]
s: 79
d: [2,80,1,8,2,0,44,47,40]
s: 80
d: [2,90,1,10,3,0,44,49,43]
d: [2,92,1,8,2,0,18,90]
s: 92
d: [2,86,1,13,2,0,44,43,47]
s: 86
d: [2,108,2,9,3,0,44,76,40]
s: 108
d: [2,87,1,13,2,0,44,43,47]
s: 87
d: [2,85,1,14,3,0,44,43,40]
s: 85
d: [2,155,2,0,0,3,52,64,59,54,85]
e: 155
d: [2,167,2,12,3,1,44,54,85]
d: [2,168,2,5,2,1,19,167]
d: [2,178,2,3,1,2,21,168,64]
d: [2,179,2,0,0,3,21,178,59]
e: 179
f: 179,36,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t99_relat_1
% SZS output start Proof for relat_1__t99_relat_1
1. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(X0,X2))) [input t99_relat_1]
2. ~! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(X0,X2))) [negated conjecture 1]
14. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
16. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(k3_xboole_0(X0,X1),X2)) [input t96_relat_1]
17. ~! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(X0,X2))) [flattening 2]
18. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 14]
20. ! [X0,X1,X2] : (v1_relat_1(X2) => k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(k3_xboole_0(X0,X1),X2)) [flattening 16]
23. ? [X0,X1,X2] : ((k6_relat_1(X0,k6_relat_1(X1,X2)) != k6_relat_1(X0,X2) & r1_tarski(X0,X1)) & v1_relat_1(X2)) [ennf transformation 17]
24. ? [X0,X1,X2] : (k6_relat_1(X0,k6_relat_1(X1,X2)) != k6_relat_1(X0,X2) & r1_tarski(X0,X1) & v1_relat_1(X2)) [flattening 23]
26. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 18]
27. ! [X0,X1,X2] : (k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(k3_xboole_0(X0,X1),X2) | ~v1_relat_1(X2)) [ennf transformation 20]
31. ? [X0,X1,X2] : (k6_relat_1(X0,k6_relat_1(X1,X2)) != k6_relat_1(X0,X2) & r1_tarski(X0,X1) & v1_relat_1(X2)) => (k6_relat_1(sK0,k6_relat_1(sK1,sK2)) != k6_relat_1(sK0,sK2) & r1_tarski(sK0,sK1) & v1_relat_1(sK2)) [choice axiom]
32. k6_relat_1(sK0,k6_relat_1(sK1,sK2)) != k6_relat_1(sK0,sK2) & r1_tarski(sK0,sK1) & v1_relat_1(sK2) [skolemisation 24,31]
36. v1_relat_1(sK2) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. r1_tarski(sK0,sK1) [cnf transformation 32] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. k6_relat_1(sK0,k6_relat_1(sK1,sK2)) != k6_relat_1(sK0,sK2) [cnf transformation 32] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. k6_relat_1(X0,k6_relat_1(X1,X2)) = k6_relat_1(k3_xboole_0(X0,X1),X2) | ~v1_relat_1(X2) [cnf transformation 27] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. 1 <=> k6_relat_1(sK0,k6_relat_1(sK1,sK2)) = k6_relat_1(sK0,sK2) [avatar definition]
54. k6_relat_1(sK0,k6_relat_1(sK1,sK2)) != k6_relat_1(sK0,sK2) <- (~1) [avatar component clause 52] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
55. ~1 [avatar split clause 38,52]
57. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
59. r1_tarski(sK0,sK1) <- (2) [avatar component clause 57] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 2 [avatar split clause 37,57]
62. 3 <=> v1_relat_1(sK2) [avatar definition]
64. v1_relat_1(sK2) <- (3) [avatar component clause 62] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. 3 [avatar split clause 36,62]
85. k6_relat_1(X2,k6_relat_1(X3,X4)) = k6_relat_1(X2,X4) | ~v1_relat_1(X4) | ~r1_tarski(X2,X3) [superposition 43,40] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
167. k6_relat_1(sK0,sK2) != k6_relat_1(sK0,sK2) | ~v1_relat_1(sK2) | ~r1_tarski(sK0,sK1) <- (~1) [superposition 54,85] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
168. ~v1_relat_1(sK2) | ~r1_tarski(sK0,sK1) <- (~1) [trivial inequality removal 167] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
178. ~r1_tarski(sK0,sK1) <- (~1, 3) [subsumption resolution 168,64] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
179. $false <- (~1, 2, 3) [subsumption resolution 178,59] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
180. 1 | ~2 | ~3 [avatar contradiction clause 179]
181. $false [avatar sat refutation 55,60,65,180] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:36,thDist:-36}
% SZS output end Proof for relat_1__t99_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 115
% Active clauses: 25
% Passive clauses: 47
% Final active clauses: 25
% Final passive clauses: 16
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Duplicate literals: 12
% Trivial inequalities: 1
% Fw subsumption resolutions: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Equational tautologies: 1
% Forward subsumptions: 28
% 
% >>> Generating Inferences
% Binary resolution: 10
% Unit resulting resolution: 1
% Forward superposition: 29
% Backward superposition: 29
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

