i: [1,108,0,2,1,1,0,0] t51_boolealg
a: [3,194,0,2,1,108]
s: 194
i: [1,107,0,2,1,1,0,0] t51_boolealg
a: [3,199,0,2,1,107]
s: 199
i: [1,106,0,2,1,1,0,0] t51_boolealg
a: [3,204,0,2,1,106]
s: 204
i: [1,105,0,2,1,1,0,0] t51_boolealg
a: [3,209,0,2,1,105]
s: 209
i: [1,114,0,15,3,1,0,0] t51_boolealg
a: [3,168,0,4,1,114]
s: 168
i: [1,111,0,4,1,1,0,0] t51_boolealg
a: [3,179,0,4,1,111]
s: 179
i: [1,110,0,4,1,1,0,0] t51_boolealg
a: [3,184,0,4,1,110]
s: 184
i: [1,109,0,4,1,1,0,0] t51_boolealg
a: [3,189,0,4,1,109]
s: 189
a: [3,163,0,7,1,114]
s: 163
i: [1,159,0,2,1,0,0,2] existence_l2_lattices
a: [3,214,0,2,1,159]
s: 214
i: [1,160,0,2,1,0,0,1] existence_l3_lattices
a: [3,219,0,2,1,160]
s: 219
i: [1,115,0,4,2,0,0,255] dt_l3_lattices
s: 115
i: [1,141,0,4,1,0,0,1] existence_m1_subset_1
s: 141
d: [2,221,1,2,1,1,40,115,194]
a: [3,226,1,2,1,221]
s: 226
i: [1,123,0,8,4,0,0,255] cc4_lattices
s: 123
i: [1,124,0,8,4,0,0,255] cc4_lattices
s: 124
d: [2,222,1,2,1,1,40,115,219]
a: [3,231,1,2,1,222]
s: 231
i: [1,126,0,8,4,0,0,255] cc5_lattices
s: 126
d: [2,233,1,6,3,1,40,126,194]
d: [2,235,1,4,2,2,21,233,199]
d: [2,236,1,2,1,3,21,235,209]
a: [3,240,1,2,1,236]
s: 240
i: [1,127,0,8,4,0,0,255] cc5_lattices
s: 127
i: [1,128,0,8,4,0,0,255] cc5_lattices
s: 128
d: [2,255,1,6,3,1,40,127,194]
d: [2,257,1,4,2,2,21,255,199]
d: [2,258,1,2,1,3,21,257,209]
a: [3,262,1,2,1,258]
s: 262
i: [1,132,0,8,4,0,0,255] cc1_lattices
s: 132
d: [2,269,1,6,3,1,40,128,194]
d: [2,271,1,4,2,2,21,269,199]
d: [2,272,1,2,1,3,21,271,209]
a: [3,276,1,2,1,272]
s: 276
i: [1,133,0,8,4,0,0,255] cc1_lattices
s: 133
i: [1,134,0,8,4,0,0,255] cc1_lattices
s: 134
d: [2,299,1,6,3,1,40,132,194]
d: [2,301,1,4,2,2,21,299,204]
d: [2,302,1,2,1,3,21,301,209]
a: [3,306,1,2,1,302]
s: 306
i: [1,135,0,8,4,0,0,255] cc1_lattices
s: 135
d: [2,317,1,6,3,1,40,133,194]
d: [2,319,1,4,2,2,21,317,204]
d: [2,320,1,2,1,3,21,319,209]
a: [3,324,1,2,1,320]
s: 324
i: [1,136,0,8,4,0,0,255] cc1_lattices
s: 136
i: [1,137,0,8,4,0,0,255] cc1_lattices
s: 137
d: [2,331,1,6,3,1,40,134,194]
d: [2,333,1,4,2,2,21,331,204]
d: [2,334,1,2,1,3,21,333,209]
a: [3,338,1,2,1,334]
s: 338
i: [1,117,0,10,5,0,0,255] cc3_lattices
s: 117
d: [2,345,1,6,3,1,40,135,194]
d: [2,347,1,4,2,2,21,345,204]
d: [2,348,1,2,1,3,21,347,209]
a: [3,352,1,2,1,348]
s: 352
i: [1,119,0,12,6,0,0,255] cc6_lattices
s: 119
i: [1,121,0,18,9,0,0,255] cc2_lattices
s: 121
d: [2,359,1,6,3,1,40,136,194]
d: [2,361,1,4,2,2,21,359,204]
d: [2,362,1,2,1,3,21,361,209]
a: [3,366,1,2,1,362]
s: 366
i: [1,146,0,19,5,0,0,255] dt_k1_lattices
s: 146
d: [2,373,1,6,3,1,40,137,194]
d: [2,375,1,4,2,2,21,373,204]
d: [2,376,1,2,1,3,21,375,209]
a: [3,380,1,2,1,376]
s: 380
i: [1,142,0,20,7,0,0,255] rd1_lattices
s: 142
i: [1,143,0,21,6,0,0,1] dt_k3_lattices
s: 143
d: [2,283,2,6,3,1,40,262,124]
d: [2,285,2,4,2,2,21,283,209]
d: [2,286,2,2,1,3,21,285,194]
a: [3,290,2,2,1,286]
s: 290
i: [1,147,0,22,7,0,0,2] symmetry_r3_boolealg
s: 147
d: [2,284,2,6,3,1,40,262,123]
d: [2,292,2,4,2,2,21,284,209]
d: [2,293,2,2,1,3,21,292,194]
a: [3,297,2,2,1,293]
s: 297
i: [1,148,0,22,7,0,0,2] symmetry_r3_boolealg
s: 148
i: [1,149,0,22,7,0,0,1] symmetry_r4_boolealg
s: 149
r: 163
r: 168
a: [3,167,0,4,1,114]
s: 167
a: [3,171,0,4,1,114]
s: 171
a: [3,164,0,7,1,114]
s: 164
i: [1,150,0,22,7,0,0,1] symmetry_r4_boolealg
s: 150
i: [1,151,0,22,7,0,0,255] redefinition_r3_boolealg
s: 151
d: [2,403,1,24,6,1,40,164,149]
d: [2,404,1,20,5,2,21,403,189]
d: [2,405,1,18,4,3,21,404,194]
d: [2,406,1,16,3,4,21,405,204]
d: [2,407,1,14,2,5,21,406,209]
a: [3,411,1,7,1,407]
s: 411
d: [2,417,2,0,0,6,52,209,306,226,184,179,411,143]
e: 417
d: [2,418,2,14,5,1,40,411,143]
d: [2,420,2,10,4,2,21,418,179]
d: [2,421,2,6,3,3,21,420,184]
d: [2,422,2,4,2,4,21,421,226]
d: [2,423,2,2,1,5,21,422,306]
d: [2,424,2,0,0,6,21,423,209]
e: 424
r: 411
i: [1,152,0,22,7,0,0,255] redefinition_r3_boolealg
s: 152
a: [3,410,1,7,1,407]
s: 410
i: [1,153,0,22,7,0,0,255] redefinition_r3_boolealg
s: 153
i: [1,154,0,22,7,0,0,255] redefinition_r3_boolealg
s: 154
a: [3,415,1,7,1,407]
s: 415
i: [1,155,0,22,7,0,0,255] redefinition_r4_boolealg
s: 155
r: 164
r: 167
r: 415
s: 168
r: 168
r: 171
s: 167
a: [3,172,0,4,1,114]
s: 172
r: 167
r: 172
s: 168
s: 171
d: [2,392,1,18,6,1,40,149,168]
d: [2,393,1,14,5,2,21,392,184]
d: [2,394,1,10,4,3,21,393,189]
d: [2,395,1,8,3,4,21,394,194]
d: [2,396,1,6,2,5,21,395,204]
d: [2,397,1,4,1,6,21,396,209]
a: [3,401,1,4,1,397]
s: 401
r: 168
r: 401
s: 167
s: 164
a: [3,400,1,4,1,397]
s: 400
i: [1,156,0,22,7,0,0,255] redefinition_r4_boolealg
s: 156
d: [2,449,1,18,6,1,40,168,155]
d: [2,451,1,14,5,2,21,449,184]
d: [2,452,1,10,4,3,21,451,189]
d: [2,453,1,8,3,4,21,452,194]
d: [2,454,1,6,2,5,21,453,204]
d: [2,455,1,4,1,6,21,454,209]
a: [3,459,1,4,1,455]
s: 459
d: [2,529,2,0,0,7,52,209,204,194,189,184,167,459,156]
e: 529
d: [2,530,2,18,6,1,40,459,156]
d: [2,533,2,14,5,2,21,530,167]
d: [2,534,2,10,4,3,21,533,184]
d: [2,535,2,6,3,4,21,534,189]
d: [2,536,2,4,2,5,21,535,194]
d: [2,537,2,2,1,6,21,536,204]
d: [2,538,2,0,0,7,21,537,209]
e: 538
r: 459
i: [1,157,0,22,7,0,0,255] redefinition_r4_boolealg
s: 157
i: [1,158,0,22,7,0,0,255] redefinition_r4_boolealg
s: 158
s: 415
i: [1,144,0,23,6,0,0,255] redefinition_k3_lattices
s: 144
d: [2,427,1,24,6,1,40,155,164]
d: [2,429,1,17,5,2,21,427,410]
d: [2,430,1,13,4,3,21,429,189]
d: [2,431,1,11,3,4,21,430,194]
d: [2,432,1,9,2,5,21,431,204]
d: [2,433,1,7,1,6,21,432,209]
a: [3,437,1,7,1,433]
s: 437
r: 164
r: 171
r: 437
r: 415
s: 172
s: 163
d: [2,467,1,18,6,1,40,172,155]
d: [2,469,1,14,5,2,21,467,179]
d: [2,470,1,10,4,3,21,469,189]
d: [2,471,1,8,3,4,21,470,194]
d: [2,472,1,6,2,5,21,471,204]
d: [2,473,1,4,1,6,21,472,209]
a: [3,477,1,4,1,473]
s: 477
r: 167
r: 172
r: 400
r: 477
s: 168
s: 171
i: [1,145,0,23,6,0,0,255] commutativity_k3_lattices
s: 145
i: [1,139,0,31,9,0,0,1] t49_boolealg
s: 139
s: 401
i: [1,140,0,31,9,0,0,1] t49_boolealg
s: 140
s: 459
r: 163
r: 168
r: 459
r: 401
s: 167
s: 164
s: 400
i: [1,138,0,35,10,0,0,1] t49_boolealg
s: 138
a: [3,458,1,4,1,455]
s: 458
a: [3,476,1,4,1,473]
s: 476
d: [2,468,1,18,6,1,40,172,149]
d: [2,479,1,14,5,2,21,468,179]
d: [2,480,1,10,4,3,21,479,189]
d: [2,481,1,8,3,4,21,480,194]
d: [2,482,1,6,2,5,21,481,204]
d: [2,483,1,4,1,6,21,482,209]
a: [3,487,1,4,1,483]
s: 487
d: [2,691,2,0,0,7,52,209,204,194,179,189,171,487,149]
e: 691
d: [2,693,2,18,6,1,40,487,149]
d: [2,700,2,14,5,2,21,693,171]
d: [2,701,2,10,4,3,21,700,189]
d: [2,702,2,6,3,4,21,701,179]
d: [2,703,2,4,2,5,21,702,194]
d: [2,704,2,2,1,6,21,703,204]
d: [2,705,2,0,0,7,21,704,209]
e: 705
r: 164
r: 171
r: 476
r: 167
r: 400
r: 458
r: 487
s: 168
s: 171
s: 401
s: 163
s: 459
s: 477
d: [2,748,2,0,0,7,52,209,204,194,189,179,171,477,156]
e: 748
d: [2,749,2,18,6,1,40,477,156]
d: [2,752,2,14,5,2,21,749,171]
d: [2,753,2,10,4,3,21,752,179]
d: [2,754,2,6,3,4,21,753,189]
d: [2,755,2,4,2,5,21,754,194]
d: [2,756,2,2,1,6,21,755,204]
d: [2,757,2,0,0,7,21,756,209]
e: 757
r: 477
a: [3,486,1,4,1,483]
s: 486
d: [2,503,2,18,6,1,40,401,155]
d: [2,505,2,14,5,2,21,503,189]
d: [2,506,2,10,4,3,21,505,184]
d: [2,507,2,8,3,4,21,506,194]
d: [2,508,2,6,2,5,21,507,204]
d: [2,509,2,4,1,6,21,508,209]
a: [3,513,2,4,1,509]
s: 513
r: 163
r: 168
r: 401
r: 459
r: 513
s: 167
s: 164
s: 400
s: 415
a: [3,512,2,4,1,509]
s: 512
s: 437
d: [2,234,1,6,3,1,40,126,219]
a: [3,253,1,2,1,234]
s: 253
d: [2,594,2,18,6,1,40,477,152]
d: [2,595,2,14,5,2,21,594,179]
d: [2,596,2,10,4,3,21,595,189]
d: [2,597,2,8,3,4,21,596,194]
d: [2,598,2,6,2,5,21,597,204]
d: [2,599,2,4,1,6,21,598,209]
a: [3,603,2,4,1,599]
s: 603
r: 603
s: 476
a: [3,602,2,4,1,599]
s: 602
d: [2,300,1,6,3,1,40,132,219]
a: [3,315,1,2,1,300]
s: 315
d: [2,553,1,34,12,0,44,144,142]
d: [2,558,1,24,9,0,18,553]
d: [2,559,1,22,8,0,21,558,115]
s: 559
d: [2,652,2,18,6,1,40,459,152]
d: [2,653,2,14,5,2,21,652,184]
d: [2,654,2,10,4,3,21,653,189]
d: [2,655,2,8,3,4,21,654,194]
d: [2,656,2,6,2,5,21,655,204]
d: [2,657,2,4,1,6,21,656,209]
a: [3,661,2,4,1,657]
s: 661
r: 661
s: 458
a: [3,660,2,4,1,657]
s: 660
d: [2,555,1,34,12,0,44,142,144]
d: [2,556,1,24,9,0,18,555]
d: [2,560,1,22,8,0,21,556,115]
s: 560
d: [2,625,1,45,14,0,44,139,145]
d: [2,628,1,35,11,0,18,625]
d: [2,631,1,33,10,0,21,628,132]
d: [2,632,1,31,9,0,21,631,115]
s: 632
d: [2,428,2,24,6,1,40,155,415]
d: [2,439,2,20,5,2,21,428,189]
d: [2,440,2,13,4,3,21,439,410]
d: [2,441,2,11,3,4,21,440,194]
d: [2,442,2,9,2,5,21,441,204]
d: [2,443,2,7,1,6,21,442,209]
a: [3,447,2,7,1,443]
s: 447
r: 164
r: 167
r: 400
r: 415
r: 447
r: 458
r: 512
r: 660
r: 437
s: 168
s: 401
s: 163
s: 459
s: 513
a: [3,414,1,7,1,407]
s: 414
s: 661
a: [3,436,1,7,1,433]
s: 436
d: [2,626,1,45,14,0,44,139,145]
d: [2,627,1,35,11,0,18,626]
d: [2,633,1,33,10,0,21,627,132]
d: [2,634,1,31,9,0,21,633,115]
s: 634
d: [2,765,3,18,6,1,40,513,152]
d: [2,766,3,14,5,2,21,765,189]
d: [2,767,3,10,4,3,21,766,184]
d: [2,768,3,8,3,4,21,767,194]
d: [2,769,3,6,2,5,21,768,204]
d: [2,770,3,4,1,6,21,769,209]
a: [3,774,3,4,1,770]
s: 774
d: [2,624,1,48,14,0,40,139,147]
d: [2,629,1,38,10,0,18,624]
s: 629
a: [3,446,2,7,1,443]
s: 446
d: [2,623,1,48,14,0,40,139,151]
d: [2,630,1,38,10,0,18,623]
s: 630
d: [2,933,3,0,0,10,52,199,194,209,204,661,179,189,184,436,410,630]
e: 933
r: 163
r: 168
r: 401
r: 414
r: 436
r: 446
r: 459
r: 661
r: 513
r: 774
s: 167
s: 400
s: 164
s: 458
s: 415
s: 437
s: 512
d: [2,644,1,48,14,0,40,140,147]
d: [2,649,1,38,10,0,18,644]
s: 649
s: 660
d: [2,643,1,48,14,0,40,140,151]
d: [2,650,1,38,10,0,18,643]
s: 650
s: 447
a: [3,773,3,4,1,770]
s: 773
d: [2,562,2,24,6,1,40,437,152]
d: [2,563,2,17,5,2,21,562,410]
d: [2,564,2,13,4,3,21,563,189]
d: [2,565,2,11,3,4,21,564,194]
d: [2,566,2,9,2,5,21,565,204]
d: [2,567,2,7,1,6,21,566,209]
a: [3,571,2,7,1,567]
s: 571
d: [2,1040,3,0,0,10,52,209,204,199,194,189,184,179,602,660,571,138]
e: 1040
d: [2,1041,3,28,9,1,40,571,138]
d: [2,1045,3,24,8,2,21,1041,660]
d: [2,1046,3,20,7,3,21,1045,602]
d: [2,1047,3,16,6,4,21,1046,179]
d: [2,1048,3,12,5,5,21,1047,184]
d: [2,1049,3,8,4,6,21,1048,189]
d: [2,1050,3,6,3,7,21,1049,194]
d: [2,1051,3,4,2,8,21,1050,199]
d: [2,1052,3,2,1,9,21,1051,204]
d: [2,1053,3,0,0,10,21,1052,209]
e: 1053
r: 164
r: 171
r: 415
r: 437
r: 476
r: 486
r: 602
r: 447
r: 571
s: 172
s: 163
s: 477
s: 487
s: 414
s: 603
s: 436
d: [2,1119,3,0,0,10,52,199,194,209,204,603,179,189,184,410,436,650]
e: 1119
i: [1,113,0,11,2,1,0,0] t51_boolealg
i: [1,112,0,11,2,1,0,0] t51_boolealg
d: [2,806,3,18,6,1,40,603,151]
d: [2,808,3,14,5,2,21,806,179]
d: [2,809,3,10,4,3,21,808,189]
d: [2,810,3,8,3,4,21,809,194]
d: [2,811,3,6,2,5,21,810,204]
d: [2,812,3,4,1,6,21,811,209]
d: [2,832,3,18,6,1,40,661,151]
d: [2,834,3,14,5,2,21,832,184]
d: [2,835,3,10,4,3,21,834,189]
d: [2,836,3,8,3,4,21,835,194]
d: [2,837,3,6,2,5,21,836,204]
d: [2,838,3,4,1,6,21,837,209]
d: [2,798,2,24,6,1,40,437,156]
d: [2,858,2,17,5,2,21,798,410]
d: [2,859,2,13,4,3,21,858,189]
d: [2,860,2,11,3,4,21,859,194]
d: [2,861,2,9,2,5,21,860,204]
d: [2,862,2,7,1,6,21,861,209]
d: [2,864,1,18,6,1,40,168,155]
d: [2,866,1,14,5,2,21,864,184]
d: [2,867,1,10,4,3,21,866,189]
d: [2,868,1,8,3,4,21,867,194]
d: [2,869,1,6,2,5,21,868,204]
d: [2,870,1,4,1,6,21,869,209]
d: [2,887,2,18,6,1,40,459,152]
d: [2,888,2,14,5,2,21,887,184]
d: [2,889,2,10,4,3,21,888,189]
d: [2,890,2,8,3,4,21,889,194]
d: [2,891,2,6,2,5,21,890,204]
d: [2,892,2,4,1,6,21,891,209]
d: [2,968,1,24,6,1,40,164,155]
d: [2,970,1,17,5,2,21,968,410]
d: [2,971,1,13,4,3,21,970,189]
d: [2,972,1,11,3,4,21,971,194]
d: [2,973,1,9,2,5,21,972,204]
d: [2,974,1,7,1,6,21,973,209]
d: [2,991,2,24,6,1,40,437,152]
d: [2,992,2,17,5,2,21,991,410]
d: [2,993,2,13,4,3,21,992,189]
d: [2,994,2,11,3,4,21,993,194]
d: [2,995,2,9,2,5,21,994,204]
d: [2,996,2,7,1,6,21,995,209]
d: [2,1072,3,20,7,3,21,1045,179]
d: [2,1073,3,16,6,4,21,1072,184]
d: [2,1074,3,12,5,5,21,1073,189]
d: [2,1075,3,10,4,6,21,1074,194]
d: [2,1076,3,8,3,7,21,1075,199]
d: [2,1077,3,6,2,8,21,1076,204]
d: [2,1078,3,4,1,9,21,1077,209]
d: [2,1081,1,18,6,1,40,172,155]
d: [2,1083,1,14,5,2,21,1081,179]
d: [2,1084,1,10,4,3,21,1083,189]
d: [2,1085,1,8,3,4,21,1084,194]
d: [2,1086,1,6,2,5,21,1085,204]
d: [2,1087,1,4,1,6,21,1086,209]
d: [2,1096,2,18,6,1,40,477,152]
d: [2,1097,2,14,5,2,21,1096,179]
d: [2,1098,2,10,4,3,21,1097,189]
d: [2,1099,2,8,3,4,21,1098,194]
d: [2,1100,2,6,2,5,21,1099,204]
d: [2,1101,2,4,1,6,21,1100,209]
f: 1119,1101,1087,1078,996,974,933,892,870,862,838,812,757,538,424,302,221,105,106,107,108,109,110,111,112,113,114
% Refutation found. Thanks to Tanya!
% SZS status Theorem for boolealg__t51_boolealg
% SZS output start Proof for boolealg__t51_boolealg
1. ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (~r4_boolealg(X0,X1,k3_lattices(X0,X2,X3)) <=> (~r4_boolealg(X0,X1,X3) & ~r4_boolealg(X0,X1,X2))))))) [input t51_boolealg]
2. ~! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (~r4_boolealg(X0,X1,k3_lattices(X0,X2,X3)) <=> (~r4_boolealg(X0,X1,X3) & ~r4_boolealg(X0,X1,X2))))))) [negated conjecture 1]
24. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (r3_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2))) [input redefinition_r3_boolealg]
26. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (r4_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2))) [input redefinition_r4_boolealg]
28. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l2_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)) => m1_subset_1(k3_lattices(X0,X1,X2),u1_struct_0(X0))) [input dt_k3_lattices]
29. ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))) [input dt_l3_lattices]
32. ! [X0] : (l3_lattices(X0) => ((v10_lattices(X0) & ~v2_struct_0(X0)) => (v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)))) [input cc1_lattices]
34. ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)) <=> (r3_boolealg(X0,X1,X3) | r3_boolealg(X0,X1,X2))))))) [input t49_boolealg]
40. ! [X0] : (l3_lattices(X0) => l2_lattices(X0)) [pure predicate removal 29]
41. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r4_boolealg(X0,X1,k3_lattices(X0,X2,X3)) <~> (~r4_boolealg(X0,X1,X3) & ~r4_boolealg(X0,X1,X2))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
42. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r4_boolealg(X0,X1,k3_lattices(X0,X2,X3)) <~> (~r4_boolealg(X0,X1,X3) & ~r4_boolealg(X0,X1,X2))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) [flattening 41]
43. ! [X0] : (l2_lattices(X0) | ~l3_lattices(X0)) [ennf transformation 40]
56. ! [X0] : (((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)) | (~v10_lattices(X0) | v2_struct_0(X0))) | ~l3_lattices(X0)) [ennf transformation 32]
57. ! [X0] : ((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)) | ~v10_lattices(X0) | v2_struct_0(X0) | ~l3_lattices(X0)) [flattening 56]
58. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)) <=> (r3_boolealg(X0,X1,X3) | r3_boolealg(X0,X1,X2))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 34]
59. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)) <=> (r3_boolealg(X0,X1,X3) | r3_boolealg(X0,X1,X2))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 58]
62. ! [X0,X1,X2] : (m1_subset_1(k3_lattices(X0,X1,X2),u1_struct_0(X0)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | ~v4_lattices(X0) | v2_struct_0(X0))) [ennf transformation 28]
63. ! [X0,X1,X2] : (m1_subset_1(k3_lattices(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | ~v4_lattices(X0) | v2_struct_0(X0)) [flattening 62]
78. ! [X0,X1,X2] : ((r3_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 24]
79. ! [X0,X1,X2] : ((r3_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 78]
82. ! [X0,X1,X2] : ((r4_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 26]
83. ! [X0,X1,X2] : ((r4_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 82]
86. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((r4_boolealg(X0,X1,X3) | r4_boolealg(X0,X1,X2)) | r4_boolealg(X0,X1,k3_lattices(X0,X2,X3))) & ((~r4_boolealg(X0,X1,X3) & ~r4_boolealg(X0,X1,X2)) | ~r4_boolealg(X0,X1,k3_lattices(X0,X2,X3)))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) [nnf transformation 42]
87. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((r4_boolealg(X0,X1,X3) | r4_boolealg(X0,X1,X2) | r4_boolealg(X0,X1,k3_lattices(X0,X2,X3))) & ((~r4_boolealg(X0,X1,X3) & ~r4_boolealg(X0,X1,X2)) | ~r4_boolealg(X0,X1,k3_lattices(X0,X2,X3))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) [flattening 86]
88. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((r4_boolealg(X0,X1,X3) | r4_boolealg(X0,X1,X2) | r4_boolealg(X0,X1,k3_lattices(X0,X2,X3))) & ((~r4_boolealg(X0,X1,X3) & ~r4_boolealg(X0,X1,X2)) | ~r4_boolealg(X0,X1,k3_lattices(X0,X2,X3))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : ((r4_boolealg(sK0,X1,X3) | r4_boolealg(sK0,X1,X2) | r4_boolealg(sK0,X1,k3_lattices(sK0,X2,X3))) & ((~r4_boolealg(sK0,X1,X3) & ~r4_boolealg(sK0,X1,X2)) | ~r4_boolealg(sK0,X1,k3_lattices(sK0,X2,X3))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l3_lattices(sK0) & v17_lattices(sK0) & v10_lattices(sK0) & ~v2_struct_0(sK0)) [choice axiom]
89. ? [X1] : (? [X2] : (? [X3] : ((r4_boolealg(sK0,X1,X3) | r4_boolealg(sK0,X1,X2) | r4_boolealg(sK0,X1,k3_lattices(sK0,X2,X3))) & ((~r4_boolealg(sK0,X1,X3) & ~r4_boolealg(sK0,X1,X2)) | ~r4_boolealg(sK0,X1,k3_lattices(sK0,X2,X3))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : ((r4_boolealg(sK0,sK1,X3) | r4_boolealg(sK0,sK1,X2) | r4_boolealg(sK0,sK1,k3_lattices(sK0,X2,X3))) & ((~r4_boolealg(sK0,sK1,X3) & ~r4_boolealg(sK0,sK1,X2)) | ~r4_boolealg(sK0,sK1,k3_lattices(sK0,X2,X3))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
90. ? [X2] : (? [X3] : ((r4_boolealg(sK0,sK1,X3) | r4_boolealg(sK0,sK1,X2) | r4_boolealg(sK0,sK1,k3_lattices(sK0,X2,X3))) & ((~r4_boolealg(sK0,sK1,X3) & ~r4_boolealg(sK0,sK1,X2)) | ~r4_boolealg(sK0,sK1,k3_lattices(sK0,X2,X3))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : ((r4_boolealg(sK0,sK1,X3) | r4_boolealg(sK0,sK1,sK2) | r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,X3))) & ((~r4_boolealg(sK0,sK1,X3) & ~r4_boolealg(sK0,sK1,sK2)) | ~r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,X3))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
91. ? [X3] : ((r4_boolealg(sK0,sK1,X3) | r4_boolealg(sK0,sK1,sK2) | r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,X3))) & ((~r4_boolealg(sK0,sK1,X3) & ~r4_boolealg(sK0,sK1,sK2)) | ~r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,X3))) & m1_subset_1(X3,u1_struct_0(sK0))) => ((r4_boolealg(sK0,sK1,sK3) | r4_boolealg(sK0,sK1,sK2) | r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3))) & ((~r4_boolealg(sK0,sK1,sK3) & ~r4_boolealg(sK0,sK1,sK2)) | ~r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3))) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
92. ((((r4_boolealg(sK0,sK1,sK3) | r4_boolealg(sK0,sK1,sK2) | r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3))) & ((~r4_boolealg(sK0,sK1,sK3) & ~r4_boolealg(sK0,sK1,sK2)) | ~r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3))) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l3_lattices(sK0) & v17_lattices(sK0) & v10_lattices(sK0) & ~v2_struct_0(sK0) [skolemisation 87,91,90,89,88]
93. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)) | (~r3_boolealg(X0,X1,X3) & ~r3_boolealg(X0,X1,X2))) & ((r3_boolealg(X0,X1,X3) | r3_boolealg(X0,X1,X2)) | ~r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [nnf transformation 59]
94. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)) | (~r3_boolealg(X0,X1,X3) & ~r3_boolealg(X0,X1,X2))) & (r3_boolealg(X0,X1,X3) | r3_boolealg(X0,X1,X2) | ~r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 93]
97. ! [X0,X1,X2] : (((r3_boolealg(X0,X1,X2) | ~r2_boolealg(X0,X1,X2)) & (r2_boolealg(X0,X1,X2) | ~r3_boolealg(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [nnf transformation 79]
99. ! [X0,X1,X2] : (((r4_boolealg(X0,X1,X2) | ~r2_boolealg(X0,X1,X2)) & (r2_boolealg(X0,X1,X2) | ~r4_boolealg(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [nnf transformation 83]
105. ~v2_struct_0(sK0) [cnf transformation 92] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
106. v10_lattices(sK0) [cnf transformation 92] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
107. v17_lattices(sK0) [cnf transformation 92] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
108. l3_lattices(sK0) [cnf transformation 92] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
109. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 92] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
110. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 92] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
111. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 92] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
112. ~r4_boolealg(sK0,sK1,sK2) | ~r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) [cnf transformation 92] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
113. ~r4_boolealg(sK0,sK1,sK3) | ~r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) [cnf transformation 92] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
114. r4_boolealg(sK0,sK1,sK3) | r4_boolealg(sK0,sK1,sK2) | r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) [cnf transformation 92] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
115. ~l3_lattices(X0) | l2_lattices(X0) [cnf transformation 43] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
132. ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | v4_lattices(X0) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
138. ~r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)) | r3_boolealg(X0,X1,X2) | r3_boolealg(X0,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 94] {a:0,w:35,wCS:14000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
139. r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)) | ~r3_boolealg(X0,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 94] {a:0,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
140. r3_boolealg(X0,X1,k3_lattices(X0,X2,X3)) | ~r3_boolealg(X0,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 94] {a:0,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
143. m1_subset_1(k3_lattices(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | ~v4_lattices(X0) | v2_struct_0(X0) [cnf transformation 63] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
151. ~r3_boolealg(X0,X1,X2) | r2_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 97] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
152. ~r2_boolealg(X0,X1,X2) | r3_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 97] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
155. ~r4_boolealg(X0,X1,X2) | r2_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 99] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
156. ~r2_boolealg(X0,X1,X2) | r4_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 99] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
162. 1 <=> r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) [avatar definition]
164. r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) <- (1) [avatar component clause 162] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
166. 2 <=> r4_boolealg(sK0,sK1,sK2) [avatar definition]
167. ~r4_boolealg(sK0,sK1,sK2) <- (~2) [avatar component clause 166] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
168. r4_boolealg(sK0,sK1,sK2) <- (2) [avatar component clause 166] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
170. 3 <=> r4_boolealg(sK0,sK1,sK3) [avatar definition]
171. ~r4_boolealg(sK0,sK1,sK3) <- (~3) [avatar component clause 170] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
172. r4_boolealg(sK0,sK1,sK3) <- (3) [avatar component clause 170] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
173. 1 | 2 | 3 [avatar split clause 114,170,166,162]
174. ~1 | ~3 [avatar split clause 113,170,162]
175. ~1 | ~2 [avatar split clause 112,166,162]
177. 4 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
179. m1_subset_1(sK3,u1_struct_0(sK0)) <- (4) [avatar component clause 177] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
180. 4 [avatar split clause 111,177]
182. 5 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
184. m1_subset_1(sK2,u1_struct_0(sK0)) <- (5) [avatar component clause 182] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
185. 5 [avatar split clause 110,182]
187. 6 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
189. m1_subset_1(sK1,u1_struct_0(sK0)) <- (6) [avatar component clause 187] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
190. 6 [avatar split clause 109,187]
192. 7 <=> l3_lattices(sK0) [avatar definition]
194. l3_lattices(sK0) <- (7) [avatar component clause 192] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
195. 7 [avatar split clause 108,192]
197. 8 <=> v17_lattices(sK0) [avatar definition]
199. v17_lattices(sK0) <- (8) [avatar component clause 197] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
200. 8 [avatar split clause 107,197]
202. 9 <=> v10_lattices(sK0) [avatar definition]
204. v10_lattices(sK0) <- (9) [avatar component clause 202] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
205. 9 [avatar split clause 106,202]
207. 10 <=> v2_struct_0(sK0) [avatar definition]
209. ~v2_struct_0(sK0) <- (~10) [avatar component clause 207] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
210. ~10 [avatar split clause 105,207]
221. l2_lattices(sK0) <- (7) [resolution 115,194] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
224. 13 <=> l2_lattices(sK0) [avatar definition]
226. l2_lattices(sK0) <- (13) [avatar component clause 224] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
227. 13 | ~7 [avatar split clause 221,192,224]
299. ~v10_lattices(sK0) | v2_struct_0(sK0) | v4_lattices(sK0) <- (7) [resolution 132,194] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
301. v2_struct_0(sK0) | v4_lattices(sK0) <- (7, 9) [subsumption resolution 299,204] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
302. v4_lattices(sK0) <- (7, 9, ~10) [subsumption resolution 301,209] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
304. 25 <=> v4_lattices(sK0) [avatar definition]
306. v4_lattices(sK0) <- (25) [avatar component clause 304] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
307. 25 | ~7 | ~9 | 10 [avatar split clause 302,207,202,192,304]
409. 39 <=> m1_subset_1(k3_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) [avatar definition]
410. m1_subset_1(k3_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) <- (39) [avatar component clause 409] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
411. ~m1_subset_1(k3_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) <- (~39) [avatar component clause 409] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
418. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_lattices(sK0) | ~v4_lattices(sK0) | v2_struct_0(sK0) <- (~39) [resolution 411,143] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
420. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_lattices(sK0) | ~v4_lattices(sK0) | v2_struct_0(sK0) <- (4, ~39) [subsumption resolution 418,179] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
421. ~l2_lattices(sK0) | ~v4_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, ~39) [subsumption resolution 420,184] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
422. ~v4_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 13, ~39) [subsumption resolution 421,226] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
423. v2_struct_0(sK0) <- (4, 5, 13, 25, ~39) [subsumption resolution 422,306] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
424. $false <- (4, 5, ~10, 13, 25, ~39) [subsumption resolution 423,209] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
425. ~4 | ~5 | 10 | ~13 | ~25 | 39 [avatar contradiction clause 424]
435. 41 <=> r2_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) [avatar definition]
436. ~r2_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) <- (~41) [avatar component clause 435] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
437. r2_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) <- (41) [avatar component clause 435] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
457. 43 <=> r2_boolealg(sK0,sK1,sK2) [avatar definition]
459. r2_boolealg(sK0,sK1,sK2) <- (43) [avatar component clause 457] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
475. 44 <=> r2_boolealg(sK0,sK1,sK3) [avatar definition]
477. r2_boolealg(sK0,sK1,sK3) <- (44) [avatar component clause 475] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
530. r4_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (43) [resolution 459,156] {a:2,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
533. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~2, 43) [subsumption resolution 530,167] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
534. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~2, 5, 43) [subsumption resolution 533,184] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
535. ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~2, 5, 6, 43) [subsumption resolution 534,189] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
536. ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~2, 5, 6, 7, 43) [subsumption resolution 535,194] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
537. v2_struct_0(sK0) <- (~2, 5, 6, 7, 9, 43) [subsumption resolution 536,204] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
538. $false <- (~2, 5, 6, 7, 9, ~10, 43) [subsumption resolution 537,209] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
539. 2 | ~5 | ~6 | ~7 | ~9 | 10 | ~43 [avatar contradiction clause 538]
569. 47 <=> r3_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) [avatar definition]
571. r3_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) <- (47) [avatar component clause 569] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
601. 48 <=> r3_boolealg(sK0,sK1,sK3) [avatar definition]
603. r3_boolealg(sK0,sK1,sK3) <- (48) [avatar component clause 601] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
623. ~r3_boolealg(X0,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | r2_boolealg(X0,X1,k3_lattices(X0,X2,X3)) | ~m1_subset_1(k3_lattices(X0,X2,X3),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [resolution 139,151] {a:1,w:48,wCS:19200,sine:1,thAx:0,allAx:2,thDist:-2}
630. ~m1_subset_1(k3_lattices(X0,X2,X3),u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | r2_boolealg(X0,X1,k3_lattices(X0,X2,X3)) | ~r3_boolealg(X0,X1,X2) [duplicate literal removal 623] {a:1,w:38,wCS:15200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
643. ~r3_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | r2_boolealg(X0,X1,k3_lattices(X0,X3,X2)) | ~m1_subset_1(k3_lattices(X0,X3,X2),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [resolution 140,151] {a:1,w:48,wCS:19200,sine:1,thAx:0,allAx:2,thDist:-2}
650. ~m1_subset_1(k3_lattices(X0,X3,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v17_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | r2_boolealg(X0,X1,k3_lattices(X0,X3,X2)) | ~r3_boolealg(X0,X1,X2) [duplicate literal removal 643] {a:1,w:38,wCS:15200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
659. 49 <=> r3_boolealg(sK0,sK1,sK2) [avatar definition]
660. ~r3_boolealg(sK0,sK1,sK2) <- (~49) [avatar component clause 659] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
661. r3_boolealg(sK0,sK1,sK2) <- (49) [avatar component clause 659] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
749. r4_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (44) [resolution 477,156] {a:2,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
752. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~3, 44) [subsumption resolution 749,171] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
753. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~3, 4, 44) [subsumption resolution 752,179] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
754. ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~3, 4, 6, 44) [subsumption resolution 753,189] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
755. ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~3, 4, 6, 7, 44) [subsumption resolution 754,194] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
756. v2_struct_0(sK0) <- (~3, 4, 6, 7, 9, 44) [subsumption resolution 755,204] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
757. $false <- (~3, 4, 6, 7, 9, ~10, 44) [subsumption resolution 756,209] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
758. 3 | ~4 | ~6 | ~7 | ~9 | 10 | ~44 [avatar contradiction clause 757]
798. r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~m1_subset_1(k3_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (41) [resolution 437,156] {a:2,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
806. r2_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (48) [resolution 603,151] {a:3,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
808. r2_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 48) [subsumption resolution 806,179] {a:3,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
809. r2_boolealg(sK0,sK1,sK3) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 6, 48) [subsumption resolution 808,189] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
810. r2_boolealg(sK0,sK1,sK3) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 6, 7, 48) [subsumption resolution 809,194] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
811. r2_boolealg(sK0,sK1,sK3) | v2_struct_0(sK0) <- (4, 6, 7, 9, 48) [subsumption resolution 810,204] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
812. r2_boolealg(sK0,sK1,sK3) <- (4, 6, 7, 9, ~10, 48) [subsumption resolution 811,209] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
813. 44 | ~4 | ~6 | ~7 | ~9 | 10 | ~48 [avatar split clause 812,601,207,202,192,187,177,475]
832. r2_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (49) [resolution 661,151] {a:3,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
834. r2_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 49) [subsumption resolution 832,184] {a:3,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
835. r2_boolealg(sK0,sK1,sK2) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 6, 49) [subsumption resolution 834,189] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
836. r2_boolealg(sK0,sK1,sK2) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 6, 7, 49) [subsumption resolution 835,194] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
837. r2_boolealg(sK0,sK1,sK2) | v2_struct_0(sK0) <- (5, 6, 7, 9, 49) [subsumption resolution 836,204] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
838. r2_boolealg(sK0,sK1,sK2) <- (5, 6, 7, 9, ~10, 49) [subsumption resolution 837,209] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
839. 43 | ~5 | ~6 | ~7 | ~9 | 10 | ~49 [avatar split clause 838,659,207,202,192,187,182,457]
858. r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (39, 41) [subsumption resolution 798,410] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
859. r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (6, 39, 41) [subsumption resolution 858,189] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
860. r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (6, 7, 39, 41) [subsumption resolution 859,194] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
861. r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | v2_struct_0(sK0) <- (6, 7, 9, 39, 41) [subsumption resolution 860,204] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
862. r4_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) <- (6, 7, 9, ~10, 39, 41) [subsumption resolution 861,209] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
863. 1 | ~6 | ~7 | ~9 | 10 | ~39 | ~41 [avatar split clause 862,435,409,207,202,192,187,162]
864. r2_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (2) [resolution 168,155] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
866. r2_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (2, 5) [subsumption resolution 864,184] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
867. r2_boolealg(sK0,sK1,sK2) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (2, 5, 6) [subsumption resolution 866,189] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
868. r2_boolealg(sK0,sK1,sK2) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (2, 5, 6, 7) [subsumption resolution 867,194] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
869. r2_boolealg(sK0,sK1,sK2) | v2_struct_0(sK0) <- (2, 5, 6, 7, 9) [subsumption resolution 868,204] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
870. r2_boolealg(sK0,sK1,sK2) <- (2, 5, 6, 7, 9, ~10) [subsumption resolution 869,209] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
871. 43 | ~2 | ~5 | ~6 | ~7 | ~9 | 10 [avatar split clause 870,207,202,192,187,182,166,457]
887. r3_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (43) [resolution 459,152] {a:2,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
888. r3_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 43) [subsumption resolution 887,184] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
889. r3_boolealg(sK0,sK1,sK2) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 6, 43) [subsumption resolution 888,189] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
890. r3_boolealg(sK0,sK1,sK2) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 6, 7, 43) [subsumption resolution 889,194] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
891. r3_boolealg(sK0,sK1,sK2) | v2_struct_0(sK0) <- (5, 6, 7, 9, 43) [subsumption resolution 890,204] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
892. r3_boolealg(sK0,sK1,sK2) <- (5, 6, 7, 9, ~10, 43) [subsumption resolution 891,209] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
893. 49 | ~5 | ~6 | ~7 | ~9 | 10 | ~43 [avatar split clause 892,457,207,202,192,187,182,659]
933. $false <- (4, 5, 6, 7, 8, 9, ~10, 39, ~41, 49) [unit resulting resolution 199,194,209,204,661,179,189,184,436,410,630] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:56,thDist:-56}
943. ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 10 | ~39 | 41 | ~49 [avatar contradiction clause 933]
968. r2_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~m1_subset_1(k3_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1) [resolution 164,155] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
970. r2_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1, 39) [subsumption resolution 968,410] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
971. r2_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1, 6, 39) [subsumption resolution 970,189] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
972. r2_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1, 6, 7, 39) [subsumption resolution 971,194] {a:1,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
973. r2_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | v2_struct_0(sK0) <- (1, 6, 7, 9, 39) [subsumption resolution 972,204] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
974. r2_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) <- (1, 6, 7, 9, ~10, 39) [subsumption resolution 973,209] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
975. 41 | ~1 | ~6 | ~7 | ~9 | 10 | ~39 [avatar split clause 974,409,207,202,192,187,162,435]
991. r3_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~m1_subset_1(k3_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (41) [resolution 437,152] {a:2,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
992. r3_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (39, 41) [subsumption resolution 991,410] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
993. r3_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (6, 39, 41) [subsumption resolution 992,189] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
994. r3_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (6, 7, 39, 41) [subsumption resolution 993,194] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
995. r3_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) | v2_struct_0(sK0) <- (6, 7, 9, 39, 41) [subsumption resolution 994,204] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
996. r3_boolealg(sK0,sK1,k3_lattices(sK0,sK2,sK3)) <- (6, 7, 9, ~10, 39, 41) [subsumption resolution 995,209] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
997. 47 | ~6 | ~7 | ~9 | 10 | ~39 | ~41 [avatar split clause 996,435,409,207,202,192,187,569]
1041. r3_boolealg(sK0,sK1,sK2) | r3_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v17_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (47) [resolution 571,138] {a:3,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1045. r3_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v17_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (47, ~49) [subsumption resolution 1041,660] {a:3,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1072. r3_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v17_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 47, ~49) [subsumption resolution 1045,179] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1073. r3_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v17_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 47, ~49) [subsumption resolution 1072,184] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1074. r3_boolealg(sK0,sK1,sK3) | ~l3_lattices(sK0) | ~v17_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 6, 47, ~49) [subsumption resolution 1073,189] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1075. r3_boolealg(sK0,sK1,sK3) | ~v17_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 6, 7, 47, ~49) [subsumption resolution 1074,194] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1076. r3_boolealg(sK0,sK1,sK3) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 6, 7, 8, 47, ~49) [subsumption resolution 1075,199] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1077. r3_boolealg(sK0,sK1,sK3) | v2_struct_0(sK0) <- (4, 5, 6, 7, 8, 9, 47, ~49) [subsumption resolution 1076,204] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1078. r3_boolealg(sK0,sK1,sK3) <- (4, 5, 6, 7, 8, 9, ~10, 47, ~49) [subsumption resolution 1077,209] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1079. 48 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 10 | ~47 | 49 [avatar split clause 1078,659,569,207,202,197,192,187,182,177,601]
1081. r2_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (3) [resolution 172,155] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1083. r2_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (3, 4) [subsumption resolution 1081,179] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1084. r2_boolealg(sK0,sK1,sK3) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (3, 4, 6) [subsumption resolution 1083,189] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1085. r2_boolealg(sK0,sK1,sK3) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (3, 4, 6, 7) [subsumption resolution 1084,194] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1086. r2_boolealg(sK0,sK1,sK3) | v2_struct_0(sK0) <- (3, 4, 6, 7, 9) [subsumption resolution 1085,204] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1087. r2_boolealg(sK0,sK1,sK3) <- (3, 4, 6, 7, 9, ~10) [subsumption resolution 1086,209] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1088. 44 | ~3 | ~4 | ~6 | ~7 | ~9 | 10 [avatar split clause 1087,207,202,192,187,177,170,475]
1096. r3_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (44) [resolution 477,152] {a:2,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1097. r3_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 44) [subsumption resolution 1096,179] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1098. r3_boolealg(sK0,sK1,sK3) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 6, 44) [subsumption resolution 1097,189] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1099. r3_boolealg(sK0,sK1,sK3) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 6, 7, 44) [subsumption resolution 1098,194] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1100. r3_boolealg(sK0,sK1,sK3) | v2_struct_0(sK0) <- (4, 6, 7, 9, 44) [subsumption resolution 1099,204] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1101. r3_boolealg(sK0,sK1,sK3) <- (4, 6, 7, 9, ~10, 44) [subsumption resolution 1100,209] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1102. 48 | ~4 | ~6 | ~7 | ~9 | 10 | ~44 [avatar split clause 1101,475,207,202,192,187,177,601]
1119. $false <- (4, 5, 6, 7, 8, 9, ~10, 39, ~41, 48) [unit resulting resolution 199,194,209,204,603,179,189,184,410,436,650] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:56,thDist:-56}
1120. ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 10 | ~39 | 41 | ~48 [avatar contradiction clause 1119]
1121. $false [avatar sat refutation 173,174,175,180,185,190,195,200,205,210,227,307,425,539,758,813,839,863,871,893,943,975,997,1079,1088,1102,1120] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1511,thDist:-1511}
% SZS output end Proof for boolealg__t51_boolealg
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 33
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 3
% 
% >>> Saturation
% Initial clauses: 56
% Generated clauses: 826
% Active clauses: 141
% Passive clauses: 174
% Final active clauses: 73
% Final passive clauses: 17
% Main loop iterations started: 150
% 
% >>> Simplifying Inferences
% Duplicate literals: 144
% Fw subsumption resolutions: 436
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 53
% 
% >>> Generating Inferences
% Binary resolution: 128
% Unit resulting resolution: 7
% Forward superposition: 19
% Backward superposition: 9
% 
% >>> AVATAR
% Split clauses: 13
% Split components: 36
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 113
% SAT solver unit clauses: 9
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 5500
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% splitting model update: 0.001 s
% forward subsumption: 0.001 s
% other: 0.007 s ( own 0.005 s ) 

