i: [1,70,0,9,3,1,0,0] t33_mcart_1
a: [3,103,0,3,1,70]
s: 103
i: [1,69,0,6,1,1,0,0] t33_mcart_1
a: [3,116,0,6,1,69]
s: 116
i: [1,71,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,126,0,2,1,71]
s: 126
i: [1,72,0,2,1,0,0,1] fc1_xboole_0
a: [3,131,0,2,1,72]
s: 131
i: [1,92,0,2,1,0,0,2] rc2_xboole_0
a: [3,141,0,2,1,92]
s: 141
i: [1,93,0,2,1,0,0,2] rc1_xboole_0
a: [3,146,0,2,1,93]
s: 146
i: [1,95,0,2,1,0,0,255] rc1_relat_1
a: [3,151,0,2,1,95]
s: 151
i: [1,94,0,2,1,0,0,255] rc1_relat_1
a: [3,156,0,2,1,94]
s: 156
i: [1,96,0,2,1,0,0,255] rc2_relat_1
a: [3,161,0,2,1,96]
s: 161
i: [1,68,0,9,1,1,0,0] t33_mcart_1
a: [3,121,0,9,1,68]
s: 121
i: [1,73,0,3,1,0,0,255] d2_xboole_0
a: [3,136,0,3,1,73]
s: 136
r: 126
i: [1,74,0,4,2,0,0,255] cc1_relat_1
s: 74
i: [1,89,0,4,1,0,0,255] existence_m1_subset_1
s: 89
d: [2,164,1,2,1,1,40,74,146]
a: [3,169,1,2,1,164]
s: 169
i: [1,75,0,5,2,0,0,1] t6_boole
s: 75
i: [1,81,0,5,2,0,0,255] t7_boole
s: 81
d: [2,165,1,2,1,1,40,74,131]
a: [3,174,1,2,1,165]
s: 174
i: [1,78,0,6,2,0,0,255] t1_subset
s: 78
i: [1,90,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 90
i: [1,97,0,6,1,0,0,1] t31_mcart_1
s: 97
i: [1,98,0,6,1,0,0,1] t31_mcart_1
s: 98
i: [1,99,0,6,1,0,0,1] t31_mcart_1
s: 99
i: [1,80,0,7,3,0,0,1] t8_boole
s: 80
d: [2,176,1,7,3,0,44,75,75]
s: 176
i: [1,79,0,8,3,0,0,255] t2_subset
s: 79
i: [1,91,0,11,4,0,0,1] fc11_subset_1
s: 91
d: [2,183,1,11,4,1,44,91,121]
a: [3,202,1,5,1,183]
s: 202
i: [1,82,0,15,4,0,0,1] t31_mcart_1
s: 82
d: [2,182,1,6,2,0,40,79,89]
s: 182
i: [1,86,0,21,5,0,0,1] t32_mcart_1
s: 86
i: [1,87,0,21,5,0,0,1] t32_mcart_1
s: 87
d: [2,211,1,21,5,1,44,86,121]
a: [3,224,1,3,1,211]
s: 224
r: 121
r: 103
d: [2,251,0,9,1,2,23,121,224]
d: [2,253,0,6,1,2,22,251,99]
d: [2,254,0,0,0,3,21,253,116]
e: 254
r: 224
s: 103
a: [3,223,1,3,1,211]
s: 223
s: 121
i: [1,88,0,21,5,0,0,1] t32_mcart_1
s: 88
a: [3,228,1,3,1,211]
s: 228
r: 121
d: [2,281,0,9,1,2,23,121,228]
d: [2,282,0,6,1,2,22,281,98]
d: [2,283,0,0,0,3,21,282,116]
e: 283
r: 228
s: 121
a: [3,227,1,3,1,211]
s: 227
a: [3,232,1,3,1,211]
s: 232
r: 121
d: [2,299,0,9,1,2,23,121,232]
d: [2,300,0,6,1,2,22,299,97]
d: [2,301,0,0,0,3,21,300,116]
e: 301
r: 232
s: 121
a: [3,231,1,3,1,211]
s: 231
d: [2,320,2,0,0,5,52,103,223,227,121,231,86]
e: 320
r: 103
a: [3,102,0,3,1,70]
s: 102
r: 121
a: [3,107,0,3,1,70]
s: 107
d: [2,322,0,9,1,2,23,121,102]
a: [3,327,0,9,1,322]
s: 327
a: [3,235,1,12,2,211]
s: 235
d: [2,323,2,3,1,2,44,223,102]
a: [3,332,2,3,1,323]
s: 332
d: [2,356,3,0,0,5,52,107,231,227,327,332,87]
e: 356
r: 107
a: [3,106,0,3,1,70]
s: 106
r: 327
a: [3,111,0,3,1,70]
s: 111
d: [2,358,0,9,1,2,23,327,106]
a: [3,363,0,9,1,358]
s: 363
d: [2,237,1,21,5,1,44,87,121]
a: [3,249,1,12,2,237]
s: 249
d: [2,359,2,3,1,2,44,227,106]
a: [3,368,2,3,1,359]
s: 368
d: [2,392,3,0,0,5,52,111,332,231,363,368,88]
e: 392
f: 392,359,358,356,323,322,320,301,283,254,68,69,70
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t33_mcart_1
% SZS output start Proof for mcart_1__t33_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5) => ((X2 = X5 & X1 = X4 & X0 = X3) | k3_zfmisc_1(X0,X1,X2) = k1_xboole_0)) [input t33_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5) => ((X2 = X5 & X1 = X4 & X0 = X3) | k3_zfmisc_1(X0,X1,X2) = k1_xboole_0)) [negated conjecture 1]
21. ! [X0] : ! [X1] : ! [X2] : ((k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) <=> k3_zfmisc_1(X0,X1,X2) != k1_xboole_0) [input t31_mcart_1]
22. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5) => ((X2 = X5 & X1 = X4 & X0 = X3) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t32_mcart_1]
26. ~! [X0,X1,X2,X3,X4,X5] : (k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5) => ((X2 = X5 & X1 = X4 & X0 = X3) | k3_zfmisc_1(X0,X1,X2) = k1_xboole_0)) [flattening 2]
31. ! [X0,X1,X2] : ((k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) <=> k3_zfmisc_1(X0,X1,X2) != k1_xboole_0) [flattening 21]
32. ! [X0,X1,X2,X3,X4,X5] : (k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5) => ((X2 = X5 & X1 = X4 & X0 = X3) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 22]
36. ? [X0,X1,X2,X3,X4,X5] : (((X2 != X5 | X1 != X4 | X0 != X3) & k3_zfmisc_1(X0,X1,X2) != k1_xboole_0) & k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5)) [ennf transformation 26]
37. ? [X0,X1,X2,X3,X4,X5] : ((X2 != X5 | X1 != X4 | X0 != X3) & k3_zfmisc_1(X0,X1,X2) != k1_xboole_0 & k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5)) [flattening 36]
49. ! [X0,X1,X2,X3,X4,X5] : (((X2 = X5 & X1 = X4 & X0 = X3) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0) | k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5)) [ennf transformation 32]
50. ! [X0,X1,X2,X3,X4,X5] : ((X2 = X5 & X1 = X4 & X0 = X3) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5)) [flattening 49]
54. ? [X0,X1,X2,X3,X4,X5] : ((X2 != X5 | X1 != X4 | X0 != X3) & k3_zfmisc_1(X0,X1,X2) != k1_xboole_0 & k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5)) => ((sK2 != sK5 | sK1 != sK4 | sK0 != sK3) & k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) & k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK3,sK4,sK5)) [choice axiom]
55. (sK2 != sK5 | sK1 != sK4 | sK0 != sK3) & k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) & k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK3,sK4,sK5) [skolemisation 37,54]
56. ! [X0,X1,X2] : (((k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) | k3_zfmisc_1(X0,X1,X2) = k1_xboole_0) & (k3_zfmisc_1(X0,X1,X2) != k1_xboole_0 | (k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0))) [nnf transformation 31]
57. ! [X0,X1,X2] : (((k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) | k3_zfmisc_1(X0,X1,X2) = k1_xboole_0) & (k3_zfmisc_1(X0,X1,X2) != k1_xboole_0 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 56]
68. k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK3,sK4,sK5) [cnf transformation 55] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) [cnf transformation 55] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. sK2 != sK5 | sK1 != sK4 | sK0 != sK3 [cnf transformation 55] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
83. k1_xboole_0 != X0 | k3_zfmisc_1(X0,X1,X2) = k1_xboole_0 [cnf transformation 57] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:1,thDist:-1}
84. k1_xboole_0 != X1 | k3_zfmisc_1(X0,X1,X2) = k1_xboole_0 [cnf transformation 57] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:1,thDist:-1}
85. k1_xboole_0 != X2 | k3_zfmisc_1(X0,X1,X2) = k1_xboole_0 [cnf transformation 57] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:1,thDist:-1}
86. k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X0 = X3 [cnf transformation 50] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X1 = X4 [cnf transformation 50] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
88. k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X2 = X5 [cnf transformation 50] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
97. k1_xboole_0 = k3_zfmisc_1(X0,X1,k1_xboole_0) [equality resolution 85] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
98. k1_xboole_0 = k3_zfmisc_1(X0,k1_xboole_0,X2) [equality resolution 84] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
99. k1_xboole_0 = k3_zfmisc_1(k1_xboole_0,X1,X2) [equality resolution 83] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. 1 <=> sK0 = sK3 [avatar definition]
102. sK0 = sK3 <- (1) [avatar component clause 101] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
103. sK0 != sK3 <- (~1) [avatar component clause 101] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
105. 2 <=> sK1 = sK4 [avatar definition]
106. sK1 = sK4 <- (2) [avatar component clause 105] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
107. sK1 != sK4 <- (~2) [avatar component clause 105] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
109. 3 <=> sK2 = sK5 [avatar definition]
111. sK2 != sK5 <- (~3) [avatar component clause 109] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
112. ~1 | ~2 | ~3 [avatar split clause 70,109,105,101]
114. 4 <=> k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) [avatar definition]
116. k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) <- (~4) [avatar component clause 114] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
117. ~4 [avatar split clause 69,114]
119. 5 <=> k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK3,sK4,sK5) [avatar definition]
121. k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK3,sK4,sK5) <- (5) [avatar component clause 119] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
122. 5 [avatar split clause 68,119]
222. 20 <=> k1_xboole_0 = sK3 [avatar definition]
223. k1_xboole_0 != sK3 <- (~20) [avatar component clause 222] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
224. k1_xboole_0 = sK3 <- (20) [avatar component clause 222] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
226. 21 <=> k1_xboole_0 = sK4 [avatar definition]
227. k1_xboole_0 != sK4 <- (~21) [avatar component clause 226] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
228. k1_xboole_0 = sK4 <- (21) [avatar component clause 226] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
230. 22 <=> k1_xboole_0 = sK5 [avatar definition]
231. k1_xboole_0 != sK5 <- (~22) [avatar component clause 230] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
232. k1_xboole_0 = sK5 <- (22) [avatar component clause 230] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
251. k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(k1_xboole_0,sK4,sK5) <- (5, 20) [backward demodulation 121,224] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
253. k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) <- (5, 20) [forward demodulation 251,99] {a:0,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
254. $false <- (~4, 5, 20) [subsumption resolution 253,116] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
255. 4 | ~5 | ~20 [avatar contradiction clause 254]
281. k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK3,k1_xboole_0,sK5) <- (5, 21) [backward demodulation 121,228] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
282. k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) <- (5, 21) [forward demodulation 281,98] {a:0,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
283. $false <- (~4, 5, 21) [subsumption resolution 282,116] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
284. 4 | ~5 | ~21 [avatar contradiction clause 283]
299. k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK3,sK4,k1_xboole_0) <- (5, 22) [backward demodulation 121,232] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
300. k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) <- (5, 22) [forward demodulation 299,97] {a:0,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
301. $false <- (~4, 5, 22) [subsumption resolution 300,116] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
302. 4 | ~5 | ~22 [avatar contradiction clause 301]
320. $false <- (~1, 5, ~20, ~21, ~22) [unit resulting resolution 103,223,227,121,231,86] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
321. 1 | ~5 | 20 | 21 | 22 [avatar contradiction clause 320]
322. k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK0,sK4,sK5) <- (1, 5) [backward demodulation 121,102] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
323. k1_xboole_0 != sK0 <- (1, ~20) [superposition 223,102] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
325. 26 <=> k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK0,sK4,sK5) [avatar definition]
327. k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK0,sK4,sK5) <- (26) [avatar component clause 325] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
328. 26 | ~1 | ~5 [avatar split clause 322,119,101,325]
330. 27 <=> k1_xboole_0 = sK0 [avatar definition]
332. k1_xboole_0 != sK0 <- (~27) [avatar component clause 330] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
333. ~27 | ~1 | 20 [avatar split clause 323,222,101,330]
356. $false <- (~2, ~21, ~22, 26, ~27) [unit resulting resolution 107,231,227,327,332,87] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:16,thDist:-16}
357. 2 | 21 | 22 | ~26 | 27 [avatar contradiction clause 356]
358. k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK0,sK1,sK5) <- (2, 26) [backward demodulation 327,106] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
359. k1_xboole_0 != sK1 <- (2, ~21) [superposition 227,106] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
361. 29 <=> k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK0,sK1,sK5) [avatar definition]
363. k3_zfmisc_1(sK0,sK1,sK2) = k3_zfmisc_1(sK0,sK1,sK5) <- (29) [avatar component clause 361] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
364. 29 | ~2 | ~26 [avatar split clause 358,325,105,361]
366. 30 <=> k1_xboole_0 = sK1 [avatar definition]
368. k1_xboole_0 != sK1 <- (~30) [avatar component clause 366] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
369. ~30 | ~2 | 21 [avatar split clause 359,226,105,366]
392. $false <- (~3, ~22, ~27, 29, ~30) [unit resulting resolution 111,332,231,363,368,88] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
393. 3 | 22 | 27 | ~29 | 30 [avatar contradiction clause 392]
394. $false [avatar sat refutation 112,117,122,255,284,302,321,328,333,357,364,369,393] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:277,thDist:-277}
% SZS output end Proof for mcart_1__t33_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 186
% Active clauses: 49
% Passive clauses: 73
% Final active clauses: 37
% Final passive clauses: 21
% Main loop iterations started: 52
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 27
% Fw demodulations: 5
% Bw demodulations: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Equational tautologies: 3
% Forward subsumptions: 27
% 
% >>> Generating Inferences
% Binary resolution: 6
% Unit resulting resolution: 3
% Forward superposition: 24
% Backward superposition: 41
% Self superposition: 1
% Equality resolution: 3
% 
% >>> AVATAR
% Split clauses: 19
% Split components: 51
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 45
% SAT solver unit clauses: 10
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 5117
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% forward subsumption index maintenance: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.005 s ( own 0.0000 s ) 

