i: [1,65,0,12,3,1,0,0] t137_zfmisc_1
a: [3,98,0,3,1,65]
s: 98
a: [3,102,0,3,1,65]
s: 102
a: [3,95,0,6,1,65]
s: 95
i: [1,66,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,110,0,2,1,66]
s: 110
i: [1,67,0,2,1,0,0,255] fc1_xboole_0
a: [3,115,0,2,1,67]
s: 115
i: [1,87,0,2,1,0,0,2] rc2_zfmisc_1
a: [3,125,0,2,1,87]
s: 125
i: [1,88,0,2,1,0,0,2] rc2_xboole_0
a: [3,130,0,2,1,88]
s: 130
i: [1,89,0,2,1,0,0,2] rc1_xboole_0
a: [3,135,0,2,1,89]
s: 135
i: [1,91,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,140,0,2,1,91]
s: 140
i: [1,90,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,145,0,2,1,90]
s: 145
i: [1,69,0,3,1,0,0,1] fc2_xboole_0
s: 69
i: [1,70,0,3,1,0,0,1] fc2_zfmisc_1
s: 70
i: [1,81,0,3,1,0,0,1] reflexivity_r1_tarski
s: 81
i: [1,68,0,3,1,0,0,255] d2_xboole_0
a: [3,120,0,3,1,68]
s: 120
r: 110
i: [1,72,0,4,2,0,0,2] cc2_zfmisc_1
s: 72
i: [1,73,0,4,2,0,0,2] cc1_zfmisc_1
s: 73
d: [2,148,1,2,1,1,40,72,135]
a: [3,153,1,2,1,148]
s: 153
i: [1,71,0,5,1,0,0,255] t1_boole
s: 71
d: [2,149,1,2,1,1,40,72,115]
a: [3,158,1,2,1,149]
s: 158
i: [1,74,0,5,2,0,0,255] t6_boole
s: 74
i: [1,78,0,5,2,0,0,1] t7_boole
s: 78
d: [2,164,1,2,1,1,40,78,98]
a: [3,168,1,2,1,164]
s: 168
i: [1,82,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 82
i: [1,84,0,6,2,0,0,1] fc4_xboole_0
s: 84
i: [1,85,0,6,2,0,0,1] fc5_xboole_0
s: 85
i: [1,86,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 86
d: [2,174,1,3,1,1,40,86,98]
a: [3,178,1,3,1,174]
s: 178
i: [1,75,0,7,2,0,0,1] l1_zfmisc_1
s: 75
d: [2,180,1,4,1,0,40,75,81]
s: 180
i: [1,76,0,7,2,0,0,1] l1_zfmisc_1
s: 76
i: [1,77,0,7,3,0,0,255] t8_boole
s: 77
d: [2,181,2,4,1,0,40,180,86]
s: 181
i: [1,83,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 83
d: [2,184,1,5,1,0,44,83,71]
s: 184
i: [1,79,0,8,2,0,0,1] t11_xboole_1
s: 79
i: [1,80,0,11,3,0,0,1] t8_xboole_1
s: 80
d: [2,205,1,7,2,1,40,80,95]
d: [2,213,1,4,1,2,21,205,102]
a: [3,217,1,4,1,213]
s: 217
d: [2,219,2,0,0,2,52,98,217,76]
e: 219
d: [2,220,2,3,1,1,40,217,76]
d: [2,222,2,0,0,2,21,220,98]
e: 222
r: 95
r: 98
r: 217
a: [3,99,0,3,1,65]
s: 99
a: [3,94,0,6,1,65]
s: 94
d: [2,185,1,5,1,0,44,83,71]
s: 185
d: [2,186,1,5,1,0,44,71,83]
s: 186
d: [2,189,1,5,1,0,44,71,83]
s: 189
d: [2,199,1,5,1,0,40,79,81]
s: 199
d: [2,231,2,3,1,0,44,199,184]
s: 231
d: [2,160,1,7,3,0,44,74,74]
s: 160
d: [2,162,1,7,2,0,44,71,74]
s: 162
d: [2,229,2,5,1,0,44,199,83]
s: 229
d: [2,202,1,8,2,0,44,79,83]
s: 202
a: [3,216,1,4,1,213]
s: 216
d: [2,265,2,0,0,2,52,99,216,75]
e: 265
d: [2,266,2,3,1,1,40,216,75]
d: [2,268,2,0,0,2,21,266,99]
e: 268
r: 99
r: 102
i: [1,64,0,9,2,1,0,0] t137_zfmisc_1
i: [1,63,0,9,2,1,0,0] t137_zfmisc_1
d: [2,254,2,4,1,1,40,202,94]
d: [2,224,1,3,1,1,40,94,79]
f: 224,266,254,222,213,63,64,65
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t137_zfmisc_1
% SZS output start Proof for zfmisc_1__t137_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1) <=> (r1_tarski(X2,X1) & r2_hidden(X0,X1))) [input t137_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1) <=> (r1_tarski(X2,X1) & r2_hidden(X0,X1))) [negated conjecture 1]
6. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
18. ! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [input l1_zfmisc_1]
24. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_xboole_0(X0,X1),X2) => r1_tarski(X0,X2)) [input t11_xboole_1]
29. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [input t8_xboole_1]
30. ~! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1) <=> (r1_tarski(X2,X1) & r2_hidden(X0,X1))) [flattening 2]
31. ! [X0,X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [flattening 18]
34. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X1),X2) => r1_tarski(X0,X2)) [flattening 24]
35. ! [X0,X1,X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [flattening 29]
38. ? [X0,X1,X2] : (r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1) <~> (r1_tarski(X2,X1) & r2_hidden(X0,X1))) [ennf transformation 30]
44. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(k2_xboole_0(X0,X1),X2)) [ennf transformation 34]
45. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | (~r1_tarski(X2,X1) | ~r1_tarski(X0,X1))) [ennf transformation 35]
46. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1)) [flattening 45]
50. ? [X0,X1,X2] : (((~r1_tarski(X2,X1) | ~r2_hidden(X0,X1)) | ~r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1)) & ((r1_tarski(X2,X1) & r2_hidden(X0,X1)) | r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1))) [nnf transformation 38]
51. ? [X0,X1,X2] : ((~r1_tarski(X2,X1) | ~r2_hidden(X0,X1) | ~r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1)) & ((r1_tarski(X2,X1) & r2_hidden(X0,X1)) | r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1))) [flattening 50]
52. ? [X0,X1,X2] : ((~r1_tarski(X2,X1) | ~r2_hidden(X0,X1) | ~r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1)) & ((r1_tarski(X2,X1) & r2_hidden(X0,X1)) | r1_tarski(k2_xboole_0(X2,k1_tarski(X0)),X1))) => ((~r1_tarski(sK2,sK1) | ~r2_hidden(sK0,sK1) | ~r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1)) & ((r1_tarski(sK2,sK1) & r2_hidden(sK0,sK1)) | r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1))) [choice axiom]
53. (~r1_tarski(sK2,sK1) | ~r2_hidden(sK0,sK1) | ~r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1)) & ((r1_tarski(sK2,sK1) & r2_hidden(sK0,sK1)) | r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1)) [skolemisation 51,52]
54. ! [X0,X1] : ((r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X1) | ~r1_tarski(k1_tarski(X0),X1))) [nnf transformation 31]
63. r2_hidden(sK0,sK1) | r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1) [cnf transformation 53] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
64. r1_tarski(sK2,sK1) | r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1) [cnf transformation 53] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. ~r1_tarski(sK2,sK1) | ~r2_hidden(sK0,sK1) | ~r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1) [cnf transformation 53] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
75. ~r1_tarski(k1_tarski(X0),X1) | r2_hidden(X0,X1) [cnf transformation 54] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
76. r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 54] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. ~r1_tarski(k2_xboole_0(X0,X1),X2) | r1_tarski(X0,X2) [cnf transformation 44] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
80. r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1) [cnf transformation 46] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
83. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 6] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
93. 1 <=> r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1) [avatar definition]
94. r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1) <- (1) [avatar component clause 93] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
95. ~r1_tarski(k2_xboole_0(sK2,k1_tarski(sK0)),sK1) <- (~1) [avatar component clause 93] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
97. 2 <=> r2_hidden(sK0,sK1) [avatar definition]
98. r2_hidden(sK0,sK1) <- (2) [avatar component clause 97] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
101. 3 <=> r1_tarski(sK2,sK1) [avatar definition]
102. r1_tarski(sK2,sK1) <- (3) [avatar component clause 101] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
104. ~1 | ~2 | ~3 [avatar split clause 65,101,97,93]
105. 1 | 3 [avatar split clause 64,101,93]
106. 1 | 2 [avatar split clause 63,97,93]
202. ~r1_tarski(k2_xboole_0(X5,X4),X6) | r1_tarski(X4,X6) [superposition 79,83] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
205. ~r1_tarski(k1_tarski(sK0),sK1) | ~r1_tarski(sK2,sK1) <- (~1) [resolution 80,95] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
213. ~r1_tarski(k1_tarski(sK0),sK1) <- (~1, 3) [subsumption resolution 205,102] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
215. 16 <=> r1_tarski(k1_tarski(sK0),sK1) [avatar definition]
216. r1_tarski(k1_tarski(sK0),sK1) <- (16) [avatar component clause 215] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
217. ~r1_tarski(k1_tarski(sK0),sK1) <- (~16) [avatar component clause 215] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
218. ~16 | 1 | ~3 [avatar split clause 213,101,93,215]
220. ~r2_hidden(sK0,sK1) <- (~16) [resolution 217,76] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
222. $false <- (2, ~16) [subsumption resolution 220,98] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
223. ~2 | 16 [avatar contradiction clause 222]
224. r1_tarski(sK2,sK1) <- (1) [resolution 94,79] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
254. r1_tarski(k1_tarski(sK0),sK1) <- (1) [resolution 202,94] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
264. 16 | ~1 [avatar split clause 254,93,215]
266. r2_hidden(sK0,sK1) <- (16) [resolution 216,75] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
270. 2 | ~16 [avatar split clause 266,215,97]
271. 3 | ~1 [avatar split clause 224,93,101]
272. $false [avatar sat refutation 104,105,106,218,223,264,270,271] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:96,thDist:-96}
% SZS output end Proof for zfmisc_1__t137_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 28
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 146
% Active clauses: 44
% Passive clauses: 68
% Final active clauses: 38
% Final passive clauses: 19
% Main loop iterations started: 49
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 19
% Forward subsumptions: 37
% 
% >>> Generating Inferences
% Binary resolution: 23
% Unit resulting resolution: 2
% Forward superposition: 40
% Backward superposition: 22
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 25
% SAT solver unit clauses: 8
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

