i: [1,47,0,2,1,1,0,0] t12_grfunc_1
a: [3,74,0,2,1,47]
s: 74
i: [1,46,0,2,1,1,0,0] t12_grfunc_1
a: [3,79,0,2,1,46]
s: 79
i: [1,45,0,2,1,1,0,0] t12_grfunc_1
a: [3,84,0,2,1,45]
s: 84
i: [1,44,0,2,1,1,0,0] t12_grfunc_1
a: [3,89,0,2,1,44]
s: 89
i: [1,43,0,2,1,1,0,0] t12_grfunc_1
a: [3,94,0,2,1,43]
s: 94
i: [1,48,0,4,1,1,0,0] t12_grfunc_1
a: [3,69,0,4,1,48]
s: 69
i: [1,62,0,2,1,0,0,1] rc1_funct_1
a: [3,99,0,2,1,62]
s: 99
i: [1,61,0,2,1,0,0,1] rc1_funct_1
a: [3,104,0,2,1,61]
s: 104
i: [1,65,0,2,1,0,0,1] rc2_funct_1
a: [3,109,0,2,1,65]
s: 109
i: [1,64,0,2,1,0,0,1] rc2_funct_1
a: [3,114,0,2,1,64]
s: 114
i: [1,63,0,2,1,0,0,1] rc2_funct_1
a: [3,119,0,2,1,63]
s: 119
i: [1,56,0,3,1,0,0,2] reflexivity_r1_tarski
s: 56
i: [1,55,0,4,1,0,0,255] existence_m1_subset_1
s: 55
i: [1,52,0,5,1,0,0,2] t17_xboole_1
s: 52
i: [1,57,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 57
i: [1,59,0,6,2,0,0,1] fc1_relat_1
s: 59
i: [1,53,0,7,2,0,0,255] t3_subset
s: 53
i: [1,54,0,7,2,0,0,255] t3_subset
s: 54
d: [2,123,1,5,1,0,40,53,55]
s: 123
i: [1,58,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 58
i: [1,49,0,8,3,0,0,255] cc2_relat_1
s: 49
d: [2,126,1,5,1,0,44,52,58]
s: 126
i: [1,60,0,8,3,0,0,1] fc1_grfunc_1
s: 60
d: [2,128,1,5,1,0,44,52,58]
s: 128
i: [1,50,0,10,4,0,0,255] cc3_funct_1
s: 50
i: [1,51,0,15,7,0,0,1] t10_grfunc_1
s: 51
d: [2,125,1,6,2,0,44,59,58]
s: 125
d: [2,127,1,6,2,0,44,59,58]
s: 127
d: [2,129,1,6,2,0,40,49,55]
s: 129
d: [2,130,1,7,3,0,40,49,54]
s: 130
d: [2,135,1,8,3,0,44,60,58]
s: 135
d: [2,136,1,8,3,0,44,60,58]
s: 136
d: [2,137,1,8,3,0,40,50,55]
s: 137
d: [2,138,1,9,4,0,40,50,54]
s: 138
d: [2,140,1,18,6,0,40,51,52]
d: [2,144,1,14,5,0,21,140,60]
d: [2,145,1,10,4,0,21,144,59]
s: 145
d: [2,160,2,0,0,4,52,94,89,74,69,145]
e: 160
d: [2,161,2,6,3,1,40,145,69]
d: [2,166,2,4,2,2,21,161,74]
d: [2,167,2,2,1,3,21,166,89]
d: [2,168,2,0,0,4,21,167,94]
e: 168
f: 168,43,44,47,48
% Refutation found. Thanks to Tanya!
% SZS status Theorem for grfunc_1__t12_grfunc_1
% SZS output start Proof for grfunc_1__t12_grfunc_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (v2_funct_1(X0) => v2_funct_1(k3_xboole_0(X0,X1))))) [input t12_grfunc_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (v2_funct_1(X0) => v2_funct_1(k3_xboole_0(X0,X1))))) [negated conjecture 1]
12. ! [X0,X1] : ((v1_funct_1(X0) & v1_relat_1(X0)) => v1_funct_1(k3_xboole_0(X0,X1))) [input fc1_grfunc_1]
13. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k3_xboole_0(X0,X1))) [input fc1_relat_1]
17. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ((v2_funct_1(X1) & r1_tarski(X0,X1)) => v2_funct_1(X0)))) [input t10_grfunc_1]
18. ! [X0] : ! [X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [input t17_xboole_1]
19. ! [X0,X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [flattening 18]
23. ? [X0] : (? [X1] : ((~v2_funct_1(k3_xboole_0(X0,X1)) & v2_funct_1(X0)) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
24. ? [X0] : (? [X1] : (~v2_funct_1(k3_xboole_0(X0,X1)) & v2_funct_1(X0) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 23]
28. ! [X0] : (! [X1] : ((v2_funct_1(X0) | (~v2_funct_1(X1) | ~r1_tarski(X0,X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 17]
29. ! [X0] : (! [X1] : (v2_funct_1(X0) | ~v2_funct_1(X1) | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 28]
30. ! [X0,X1] : (v1_relat_1(k3_xboole_0(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 13]
31. ! [X0,X1] : (v1_funct_1(k3_xboole_0(X0,X1)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 12]
32. ! [X0,X1] : (v1_funct_1(k3_xboole_0(X0,X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 31]
33. ? [X0] : (? [X1] : (~v2_funct_1(k3_xboole_0(X0,X1)) & v2_funct_1(X0) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (~v2_funct_1(k3_xboole_0(sK0,X1)) & v2_funct_1(sK0) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
34. ? [X1] : (~v2_funct_1(k3_xboole_0(sK0,X1)) & v2_funct_1(sK0) & v1_funct_1(X1) & v1_relat_1(X1)) => (~v2_funct_1(k3_xboole_0(sK0,sK1)) & v2_funct_1(sK0) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
35. (~v2_funct_1(k3_xboole_0(sK0,sK1)) & v2_funct_1(sK0) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 24,34,33]
43. v1_relat_1(sK0) [cnf transformation 35] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
44. v1_funct_1(sK0) [cnf transformation 35] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
47. v2_funct_1(sK0) [cnf transformation 35] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. ~v2_funct_1(k3_xboole_0(sK0,sK1)) [cnf transformation 35] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. ~r1_tarski(X0,X1) | ~v2_funct_1(X1) | v2_funct_1(X0) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. r1_tarski(k3_xboole_0(X0,X1),X0) [cnf transformation 19] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
59. v1_relat_1(k3_xboole_0(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. v1_funct_1(k3_xboole_0(X0,X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 32] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. 1 <=> v2_funct_1(k3_xboole_0(sK0,sK1)) [avatar definition]
69. ~v2_funct_1(k3_xboole_0(sK0,sK1)) <- (~1) [avatar component clause 67] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
70. ~1 [avatar split clause 48,67]
72. 2 <=> v2_funct_1(sK0) [avatar definition]
74. v2_funct_1(sK0) <- (2) [avatar component clause 72] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
75. 2 [avatar split clause 47,72]
87. 5 <=> v1_funct_1(sK0) [avatar definition]
89. v1_funct_1(sK0) <- (5) [avatar component clause 87] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
90. 5 [avatar split clause 44,87]
92. 6 <=> v1_relat_1(sK0) [avatar definition]
94. v1_relat_1(sK0) <- (6) [avatar component clause 92] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. 6 [avatar split clause 43,92]
140. ~v2_funct_1(X1) | v2_funct_1(k3_xboole_0(X1,X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(k3_xboole_0(X1,X2)) | ~v1_relat_1(k3_xboole_0(X1,X2)) [resolution 51,52] {a:1,w:18,wCS:7200,sine:1,thAx:0,allAx:2,thDist:-2}
144. ~v2_funct_1(X1) | v2_funct_1(k3_xboole_0(X1,X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_relat_1(k3_xboole_0(X1,X2)) [subsumption resolution 140,60] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
145. v2_funct_1(k3_xboole_0(X1,X2)) | ~v2_funct_1(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [subsumption resolution 144,59] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
161. ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [resolution 145,69] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
166. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 161,74] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
167. ~v1_relat_1(sK0) <- (~1, 2, 5) [subsumption resolution 166,89] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
168. $false <- (~1, 2, 5, 6) [subsumption resolution 167,94] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
169. 1 | ~2 | ~5 | ~6 [avatar contradiction clause 168]
170. $false [avatar sat refutation 70,75,90,95,169] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:58,thDist:-58}
% SZS output end Proof for grfunc_1__t12_grfunc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 81
% Active clauses: 32
% Passive clauses: 37
% Final active clauses: 32
% Final passive clauses: 2
% Main loop iterations started: 35
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 17
% 
% >>> Generating Inferences
% Binary resolution: 19
% Unit resulting resolution: 1
% Forward superposition: 16
% Backward superposition: 5
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 11
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

