i: [1,50,0,2,1,1,0,0] t32_funct_4
a: [3,83,0,2,1,50]
s: 83
i: [1,49,0,2,1,1,0,0] t32_funct_4
a: [3,88,0,2,1,49]
s: 88
i: [1,48,0,2,1,1,0,0] t32_funct_4
a: [3,93,0,2,1,48]
s: 93
i: [1,47,0,2,1,1,0,0] t32_funct_4
a: [3,98,0,2,1,47]
s: 98
i: [1,52,0,5,1,1,0,0] t32_funct_4
a: [3,73,0,5,1,52]
s: 73
i: [1,51,0,5,1,1,0,0] t32_funct_4
a: [3,78,0,5,1,51]
s: 78
i: [1,69,0,2,1,0,0,1] rc1_funct_1
a: [3,110,0,2,1,69]
s: 110
i: [1,68,0,2,1,0,0,1] rc1_funct_1
a: [3,115,0,2,1,68]
s: 115
i: [1,60,0,3,1,0,0,1] reflexivity_r1_tarski
s: 60
i: [1,59,0,4,1,0,0,255] existence_m1_subset_1
s: 59
i: [1,56,0,5,1,0,0,1] t7_xboole_1
s: 56
i: [1,61,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 61
i: [1,63,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 63
d: [2,118,1,5,1,1,40,63,78]
a: [3,122,1,5,1,118]
s: 122
i: [1,57,0,7,2,0,0,255] t3_subset
s: 57
i: [1,58,0,7,2,0,0,255] t3_subset
s: 58
d: [2,125,1,5,1,0,40,57,59]
s: 125
i: [1,62,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 62
d: [2,127,1,5,1,0,44,56,62]
s: 127
i: [1,53,0,8,3,0,0,255] cc2_relat_1
s: 53
i: [1,67,0,8,3,0,0,2] fc3_relat_1
s: 67
d: [2,128,1,5,1,0,44,56,62]
s: 128
i: [1,64,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,105,0,9,3,64]
s: 105
d: [2,132,1,6,2,0,40,53,59]
s: 132
i: [1,54,0,10,4,0,0,255] cc3_funct_1
s: 54
i: [1,65,0,12,5,0,0,1] dt_k1_funct_4
s: 65
d: [2,133,1,7,3,0,40,53,58]
s: 133
i: [1,66,0,12,5,0,0,1] dt_k1_funct_4
s: 66
d: [2,143,2,6,2,0,40,133,56]
s: 143
i: [1,55,0,20,6,0,0,1] t31_funct_4
s: 55
d: [2,138,1,8,3,0,40,54,59]
s: 138
d: [2,144,2,6,2,0,40,133,127]
s: 144
d: [2,139,1,9,4,0,40,54,58]
s: 139
d: [2,150,3,6,2,0,44,143,62]
s: 150
d: [2,155,1,18,6,0,44,56,55]
s: 155
d: [2,174,2,0,0,6,52,98,93,88,83,73,78,155]
e: 174
d: [2,175,2,13,5,1,40,155,78]
d: [2,178,2,8,4,2,21,175,73]
d: [2,179,2,6,3,3,21,178,83]
d: [2,180,2,4,2,4,21,179,88]
d: [2,181,2,2,1,5,21,180,93]
d: [2,182,2,0,0,6,21,181,98]
e: 182
f: 182,47,48,49,50,51,52
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t32_funct_4
% SZS output start Proof for funct_4__t32_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => r1_tarski(X0,k1_funct_4(X0,X1))))) [input t32_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => r1_tarski(X0,k1_funct_4(X0,X1))))) [negated conjecture 1]
18. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k1_funct_4(X0,X1) = k2_xboole_0(X0,X1)))) [input t31_funct_4]
20. ! [X0] : ! [X1] : r1_tarski(X0,k2_xboole_0(X0,X1)) [input t7_xboole_1]
21. ! [X0,X1] : r1_tarski(X0,k2_xboole_0(X0,X1)) [flattening 20]
25. ? [X0] : (? [X1] : ((~r1_tarski(X0,k1_funct_4(X0,X1)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
26. ? [X0] : (? [X1] : (~r1_tarski(X0,k1_funct_4(X0,X1)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 25]
30. ! [X0] : (! [X1] : ((k1_funct_4(X0,X1) = k2_xboole_0(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 18]
31. ! [X0] : (! [X1] : (k1_funct_4(X0,X1) = k2_xboole_0(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 30]
39. ? [X0] : (? [X1] : (~r1_tarski(X0,k1_funct_4(X0,X1)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (~r1_tarski(sK0,k1_funct_4(sK0,X1)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
40. ? [X1] : (~r1_tarski(sK0,k1_funct_4(sK0,X1)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) => (~r1_tarski(sK0,k1_funct_4(sK0,sK1)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
41. (~r1_tarski(sK0,k1_funct_4(sK0,sK1)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 26,40,39]
47. v1_relat_1(sK0) [cnf transformation 41] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. v1_funct_1(sK0) [cnf transformation 41] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. v1_relat_1(sK1) [cnf transformation 41] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. v1_funct_1(sK1) [cnf transformation 41] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) [cnf transformation 41] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. ~r1_tarski(sK0,k1_funct_4(sK0,sK1)) [cnf transformation 41] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. k1_funct_4(X0,X1) = k2_xboole_0(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. r1_tarski(X0,k2_xboole_0(X0,X1)) [cnf transformation 21] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. 1 <=> r1_tarski(sK0,k1_funct_4(sK0,sK1)) [avatar definition]
73. ~r1_tarski(sK0,k1_funct_4(sK0,sK1)) <- (~1) [avatar component clause 71] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
74. ~1 [avatar split clause 52,71]
76. 2 <=> r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) [avatar definition]
78. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) <- (2) [avatar component clause 76] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
79. 2 [avatar split clause 51,76]
81. 3 <=> v1_funct_1(sK1) [avatar definition]
83. v1_funct_1(sK1) <- (3) [avatar component clause 81] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
84. 3 [avatar split clause 50,81]
86. 4 <=> v1_relat_1(sK1) [avatar definition]
88. v1_relat_1(sK1) <- (4) [avatar component clause 86] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. 4 [avatar split clause 49,86]
91. 5 <=> v1_funct_1(sK0) [avatar definition]
93. v1_funct_1(sK0) <- (5) [avatar component clause 91] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
94. 5 [avatar split clause 48,91]
96. 6 <=> v1_relat_1(sK0) [avatar definition]
98. v1_relat_1(sK0) <- (6) [avatar component clause 96] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
99. 6 [avatar split clause 47,96]
155. ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | r1_tarski(X0,k1_funct_4(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [superposition 56,55] {a:1,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
175. r1_tarski(sK0,k1_funct_4(sK0,sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 155,78] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
178. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 175,73] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
179. ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3) [subsumption resolution 178,83] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
180. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 4) [subsumption resolution 179,88] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
181. ~v1_relat_1(sK0) <- (~1, 2, 3, 4, 5) [subsumption resolution 180,93] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
182. $false <- (~1, 2, 3, 4, 5, 6) [subsumption resolution 181,98] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
183. 1 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar contradiction clause 182]
188. $false [avatar sat refutation 74,79,84,89,94,99,183] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:64,thDist:-64}
% SZS output end Proof for funct_4__t32_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 19
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 97
% Active clauses: 33
% Passive clauses: 43
% Final active clauses: 33
% Final passive clauses: 8
% Main loop iterations started: 35
% 
% >>> Simplifying Inferences
% Duplicate literals: 15
% Fw subsumption resolutions: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 12
% Forward subsumptions: 16
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 1
% Forward superposition: 18
% Backward superposition: 10
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 8
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

