i: [1,43,0,2,1,1,0,0] t1_partfun3
a: [3,76,0,2,1,43]
s: 76
i: [1,42,0,2,1,1,0,0] t1_partfun3
a: [3,81,0,2,1,42]
s: 81
i: [1,41,0,2,1,1,0,0] t1_partfun3
a: [3,86,0,2,1,41]
s: 86
i: [1,44,0,5,1,1,0,0] t1_partfun3
a: [3,71,0,5,1,44]
s: 71
i: [1,45,0,8,1,1,0,0] t1_partfun3
a: [3,66,0,8,1,45]
s: 66
i: [1,59,0,2,1,0,0,1] rc1_funct_1
a: [3,91,0,2,1,59]
s: 91
i: [1,58,0,2,1,0,0,1] rc1_funct_1
a: [3,96,0,2,1,58]
s: 96
i: [1,62,0,2,1,0,0,1] rc2_funct_1
a: [3,101,0,2,1,62]
s: 101
i: [1,61,0,2,1,0,0,1] rc2_funct_1
a: [3,106,0,2,1,61]
s: 106
i: [1,60,0,2,1,0,0,1] rc2_funct_1
a: [3,111,0,2,1,60]
s: 111
i: [1,57,0,3,1,0,0,1] reflexivity_r1_tarski
s: 57
i: [1,56,0,4,1,0,0,255] existence_m1_subset_1
s: 56
i: [1,49,0,7,3,0,0,1] dt_k2_funct_1
s: 49
i: [1,50,0,7,3,0,0,1] dt_k2_funct_1
s: 50
i: [1,54,0,7,2,0,0,255] t3_subset
s: 54
i: [1,55,0,7,2,0,0,255] t3_subset
s: 55
d: [2,113,1,5,1,0,40,54,56]
s: 113
i: [1,46,0,9,4,0,0,1] fc6_funct_1
s: 46
i: [1,47,0,9,4,0,0,1] fc6_funct_1
s: 47
i: [1,48,0,9,4,0,0,1] fc6_funct_1
s: 48
i: [1,52,0,10,4,0,0,255] cc3_funct_1
s: 52
d: [2,115,1,8,3,0,40,52,56]
s: 115
i: [1,51,0,11,4,0,0,1] t43_funct_1
s: 51
d: [2,116,1,9,4,0,40,52,55]
s: 116
i: [1,53,0,18,5,0,0,1] t107_funct_1
s: 53
d: [2,136,1,28,8,0,44,53,51]
d: [2,138,1,25,7,0,21,136,48]
d: [2,139,1,22,6,0,21,138,50]
d: [2,140,1,19,5,0,21,139,49]
s: 140
d: [2,144,2,0,0,5,52,86,81,76,71,66,140]
e: 144
d: [2,147,2,14,5,1,44,66,140]
d: [2,148,2,11,4,1,19,147]
d: [2,153,2,6,3,2,21,148,71]
d: [2,154,2,4,2,3,21,153,76]
d: [2,155,2,2,1,4,21,154,81]
d: [2,156,2,0,0,5,21,155,86]
e: 156
f: 156,41,42,43,44,45
% Refutation found. Thanks to Tanya!
% SZS status Theorem for partfun3__t1_partfun3
% SZS output start Proof for partfun3__t1_partfun3
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((r1_tarski(X1,k9_xtuple_0(k2_funct_1(X0))) & v2_funct_1(X0)) => k7_relat_1(X0,k7_relat_1(k2_funct_1(X0),X1)) = X1)) [input t1_partfun3]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((r1_tarski(X1,k9_xtuple_0(k2_funct_1(X0))) & v2_funct_1(X0)) => k7_relat_1(X0,k7_relat_1(k2_funct_1(X0),X1)) = X1)) [negated conjecture 1]
5. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0)))) [input dt_k2_funct_1]
10. ! [X0] : ((v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(k2_funct_1(X0)) & v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0)))) [input fc6_funct_1]
14. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((r1_tarski(X1,k9_xtuple_0(X0)) & v2_funct_1(X0)) => k7_relat_1(k2_funct_1(X0),k7_relat_1(X0,X1)) = X1)) [input t107_funct_1]
16. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => k2_funct_1(k2_funct_1(X0)) = X0)) [input t43_funct_1]
19. ? [X0] : (? [X1] : (k7_relat_1(X0,k7_relat_1(k2_funct_1(X0),X1)) != X1 & (r1_tarski(X1,k9_xtuple_0(k2_funct_1(X0))) & v2_funct_1(X0))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
20. ? [X0] : (? [X1] : (k7_relat_1(X0,k7_relat_1(k2_funct_1(X0),X1)) != X1 & r1_tarski(X1,k9_xtuple_0(k2_funct_1(X0))) & v2_funct_1(X0)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 19]
21. ! [X0] : ((v2_funct_1(k2_funct_1(X0)) & v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | (~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 10]
22. ! [X0] : ((v2_funct_1(k2_funct_1(X0)) & v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 21]
23. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 5]
24. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 23]
25. ! [X0] : ((k2_funct_1(k2_funct_1(X0)) = X0 | ~v2_funct_1(X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 16]
26. ! [X0] : (k2_funct_1(k2_funct_1(X0)) = X0 | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 25]
29. ! [X0] : (! [X1] : (k7_relat_1(k2_funct_1(X0),k7_relat_1(X0,X1)) = X1 | (~r1_tarski(X1,k9_xtuple_0(X0)) | ~v2_funct_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 14]
30. ! [X0] : (! [X1] : (k7_relat_1(k2_funct_1(X0),k7_relat_1(X0,X1)) = X1 | ~r1_tarski(X1,k9_xtuple_0(X0)) | ~v2_funct_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 29]
31. ? [X0] : (? [X1] : (k7_relat_1(X0,k7_relat_1(k2_funct_1(X0),X1)) != X1 & r1_tarski(X1,k9_xtuple_0(k2_funct_1(X0))) & v2_funct_1(X0)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (k7_relat_1(sK0,k7_relat_1(k2_funct_1(sK0),X1)) != X1 & r1_tarski(X1,k9_xtuple_0(k2_funct_1(sK0))) & v2_funct_1(sK0)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
32. ? [X1] : (k7_relat_1(sK0,k7_relat_1(k2_funct_1(sK0),X1)) != X1 & r1_tarski(X1,k9_xtuple_0(k2_funct_1(sK0))) & v2_funct_1(sK0)) => (sK1 != k7_relat_1(sK0,k7_relat_1(k2_funct_1(sK0),sK1)) & r1_tarski(sK1,k9_xtuple_0(k2_funct_1(sK0))) & v2_funct_1(sK0)) [choice axiom]
33. (sK1 != k7_relat_1(sK0,k7_relat_1(k2_funct_1(sK0),sK1)) & r1_tarski(sK1,k9_xtuple_0(k2_funct_1(sK0))) & v2_funct_1(sK0)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 20,32,31]
41. v1_relat_1(sK0) [cnf transformation 33] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. v1_funct_1(sK0) [cnf transformation 33] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
43. v2_funct_1(sK0) [cnf transformation 33] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
44. r1_tarski(sK1,k9_xtuple_0(k2_funct_1(sK0))) [cnf transformation 33] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
45. sK1 != k7_relat_1(sK0,k7_relat_1(k2_funct_1(sK0),sK1)) [cnf transformation 33] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. v2_funct_1(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. v1_relat_1(k2_funct_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. v1_funct_1(k2_funct_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k2_funct_1(k2_funct_1(X0)) = X0 | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. k7_relat_1(k2_funct_1(X0),k7_relat_1(X0,X1)) = X1 | ~r1_tarski(X1,k9_xtuple_0(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. 1 <=> sK1 = k7_relat_1(sK0,k7_relat_1(k2_funct_1(sK0),sK1)) [avatar definition]
66. sK1 != k7_relat_1(sK0,k7_relat_1(k2_funct_1(sK0),sK1)) <- (~1) [avatar component clause 64] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
67. ~1 [avatar split clause 45,64]
69. 2 <=> r1_tarski(sK1,k9_xtuple_0(k2_funct_1(sK0))) [avatar definition]
71. r1_tarski(sK1,k9_xtuple_0(k2_funct_1(sK0))) <- (2) [avatar component clause 69] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. 2 [avatar split clause 44,69]
74. 3 <=> v2_funct_1(sK0) [avatar definition]
76. v2_funct_1(sK0) <- (3) [avatar component clause 74] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
77. 3 [avatar split clause 43,74]
79. 4 <=> v1_funct_1(sK0) [avatar definition]
81. v1_funct_1(sK0) <- (4) [avatar component clause 79] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. 4 [avatar split clause 42,79]
84. 5 <=> v1_relat_1(sK0) [avatar definition]
86. v1_relat_1(sK0) <- (5) [avatar component clause 84] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
87. 5 [avatar split clause 41,84]
136. k7_relat_1(X0,k7_relat_1(k2_funct_1(X0),X1)) = X1 | ~r1_tarski(X1,k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(k2_funct_1(X0)) | ~v1_funct_1(k2_funct_1(X0)) | ~v1_relat_1(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [superposition 53,51] {a:1,w:28,wCS:11200,sine:1,thAx:0,allAx:2,thDist:-2}
138. k7_relat_1(X0,k7_relat_1(k2_funct_1(X0),X1)) = X1 | ~r1_tarski(X1,k9_xtuple_0(k2_funct_1(X0))) | ~v1_funct_1(k2_funct_1(X0)) | ~v1_relat_1(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 136,48] {a:1,w:25,wCS:10000,sine:1,thAx:0,allAx:2,thDist:-2}
139. k7_relat_1(X0,k7_relat_1(k2_funct_1(X0),X1)) = X1 | ~r1_tarski(X1,k9_xtuple_0(k2_funct_1(X0))) | ~v1_relat_1(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 138,50] {a:1,w:22,wCS:8800,sine:1,thAx:0,allAx:2,thDist:-2}
140. k7_relat_1(X0,k7_relat_1(k2_funct_1(X0),X1)) = X1 | ~r1_tarski(X1,k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 139,49] {a:1,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
147. sK1 != sK1 | ~r1_tarski(sK1,k9_xtuple_0(k2_funct_1(sK0))) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [superposition 66,140] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
148. ~r1_tarski(sK1,k9_xtuple_0(k2_funct_1(sK0))) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [trivial inequality removal 147] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
153. ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 148,71] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
154. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3) [subsumption resolution 153,76] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
155. ~v1_relat_1(sK0) <- (~1, 2, 3, 4) [subsumption resolution 154,81] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
156. $false <- (~1, 2, 3, 4, 5) [subsumption resolution 155,86] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
157. 1 | ~2 | ~3 | ~4 | ~5 [avatar contradiction clause 156]
158. $false [avatar sat refutation 67,72,77,82,87,157] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:66,thDist:-66}
% SZS output end Proof for partfun3__t1_partfun3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 22
% Generated clauses: 63
% Active clauses: 24
% Passive clauses: 28
% Final active clauses: 24
% Final passive clauses: 2
% Main loop iterations started: 26
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 7
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 4
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 3
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 10
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

