i: [1,134,0,2,1,1,0,0] t49_funct_1
a: [3,215,0,2,1,134]
s: 215
i: [1,133,0,2,1,1,0,0] t49_funct_1
a: [3,220,0,2,1,133]
s: 220
i: [1,135,0,3,1,1,0,0] t49_funct_1
a: [3,210,0,3,1,135]
s: 210
i: [1,137,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,225,0,2,1,137]
s: 225
i: [1,138,0,2,1,0,0,2] fc1_xboole_0
a: [3,230,0,2,1,138]
s: 230
i: [1,187,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,240,0,2,1,187]
s: 240
i: [1,188,0,2,1,0,0,2] rc2_xboole_0
a: [3,245,0,2,1,188]
s: 245
i: [1,189,0,2,1,0,0,2] rc1_xboole_0
a: [3,250,0,2,1,189]
s: 250
i: [1,191,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,255,0,2,1,191]
s: 255
i: [1,190,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,260,0,2,1,190]
s: 260
i: [1,193,0,2,1,0,0,1] rc1_relat_1
a: [3,265,0,2,1,193]
s: 265
i: [1,192,0,2,1,0,0,1] rc1_relat_1
a: [3,270,0,2,1,192]
s: 270
i: [1,194,0,2,1,0,0,255] rc2_relat_1
a: [3,275,0,2,1,194]
s: 275
i: [1,196,0,2,1,0,0,1] rc1_funct_1
a: [3,280,0,2,1,196]
s: 280
i: [1,195,0,2,1,0,0,1] rc1_funct_1
a: [3,285,0,2,1,195]
s: 285
i: [1,198,0,2,1,0,0,255] rc2_funct_1
a: [3,290,0,2,1,198]
s: 290
i: [1,197,0,2,1,0,0,255] rc2_funct_1
a: [3,295,0,2,1,197]
s: 295
i: [1,136,0,9,1,1,0,0] t49_funct_1
a: [3,205,0,9,1,136]
s: 205
i: [1,140,0,3,1,0,0,255] fc2_xboole_0
s: 140
i: [1,141,0,3,1,0,0,255] fc2_zfmisc_1
s: 141
i: [1,139,0,3,1,0,0,255] d2_xboole_0
a: [3,235,0,3,1,139]
s: 235
r: 225
i: [1,142,0,4,2,0,0,255] cc2_zfmisc_1
s: 142
i: [1,143,0,4,2,0,0,255] cc1_zfmisc_1
s: 143
d: [2,298,1,2,1,1,40,142,250]
a: [3,303,1,2,1,298]
s: 303
i: [1,144,0,4,2,0,0,1] cc1_funct_1
s: 144
i: [1,145,0,4,2,0,0,1] cc1_relat_1
s: 145
d: [2,299,1,2,1,1,40,142,230]
a: [3,308,1,2,1,299]
s: 308
i: [1,169,0,4,1,0,0,255] existence_m1_subset_1
s: 169
d: [2,310,1,2,1,1,40,144,250]
a: [3,315,1,2,1,310]
s: 315
i: [1,170,0,4,1,0,0,255] fc3_xboole_0
s: 170
i: [1,171,0,4,1,0,0,2] fc1_zfmisc_1
s: 171
d: [2,311,1,2,1,1,40,144,230]
a: [3,320,1,2,1,311]
s: 320
i: [1,146,0,5,2,0,0,2] t6_boole
s: 146
d: [2,322,1,2,1,1,40,145,250]
a: [3,327,1,2,1,322]
s: 327
i: [1,147,0,5,2,0,0,2] fc10_relat_1
s: 147
i: [1,164,0,5,2,0,0,1] t7_boole
s: 164
d: [2,341,1,2,1,1,40,164,210]
a: [3,345,1,2,1,341]
s: 345
i: [1,172,0,5,1,0,0,255] fc1_funct_1
s: 172
d: [2,323,1,2,1,1,40,145,230]
a: [3,332,1,2,1,323]
s: 332
i: [1,173,0,5,1,0,0,255] fc5_relat_1
s: 173
i: [1,155,0,6,3,0,0,255] cc3_relat_1
s: 155
d: [2,348,0,4,2,0,21,155,145]
s: 348
i: [1,162,0,6,2,0,0,255] t1_subset
s: 162
d: [2,362,1,3,1,1,40,162,210]
a: [3,366,1,3,1,362]
s: 366
i: [1,175,0,6,2,0,0,1] dt_k5_relat_1
s: 175
i: [1,177,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 177
d: [2,368,1,3,1,1,40,177,210]
a: [3,372,1,3,1,368]
s: 372
i: [1,149,0,7,3,0,0,2] fc8_relat_1
s: 149
d: [2,349,1,2,1,1,40,348,250]
a: [3,355,1,2,1,349]
s: 355
i: [1,150,0,7,3,0,0,255] fc24_relat_1
s: 150
i: [1,163,0,7,3,0,0,1] t8_boole
s: 163
d: [2,350,1,2,1,1,40,348,230]
a: [3,360,1,2,1,350]
s: 360
i: [1,174,0,7,1,0,0,255] commutativity_k2_tarski
s: 174
d: [2,338,1,5,2,0,40,147,145]
s: 338
i: [1,148,0,8,2,0,0,2] rd4_relat_1
s: 148
i: [1,161,0,8,3,0,0,255] t2_subset
s: 161
d: [2,339,1,5,2,0,40,147,144]
s: 339
i: [1,178,0,8,3,0,0,1] fc17_relat_1
s: 178
d: [2,340,1,5,2,0,40,147,142]
s: 340
d: [2,382,0,6,2,0,21,178,145]
s: 382
i: [1,179,0,8,3,0,0,1] fc17_relat_1
s: 179
d: [2,351,1,5,2,0,40,348,147]
s: 351
i: [1,180,0,8,3,0,0,255] fc23_relat_1
s: 180
d: [2,380,1,6,2,0,40,161,169]
s: 380
i: [1,181,0,8,3,0,0,255] fc23_relat_1
s: 181
i: [1,182,0,8,3,0,0,1] fc8_funct_1
s: 182
d: [2,383,1,6,2,0,40,382,348]
s: 383
i: [1,183,0,8,3,0,0,1] fc8_funct_1
s: 183
d: [2,384,1,6,2,0,40,382,145]
s: 384
i: [1,184,0,8,3,0,0,1] fc16_relat_1
s: 184
i: [1,185,0,8,3,0,0,1] fc16_relat_1
s: 185
d: [2,385,1,6,2,0,40,382,144]
s: 385
i: [1,186,0,8,1,0,0,255] fc7_relat_1
s: 186
d: [2,386,1,6,2,0,40,382,142]
s: 386
i: [1,160,0,10,1,0,0,255] d5_tarski
s: 160
i: [1,165,0,11,3,0,0,2] t57_relat_1
s: 165
d: [2,388,2,6,2,0,40,380,177]
s: 388
i: [1,176,0,11,2,0,0,1] rd5_relat_1
s: 176
d: [2,334,1,7,3,0,44,146,146]
s: 334
i: [1,166,0,12,3,0,0,2] t57_relat_1
s: 166
i: [1,199,0,13,4,0,0,1] d2_funct_1
s: 199
d: [2,435,1,19,4,1,44,205,199]
a: [3,439,1,4,1,435]
s: 439
d: [2,453,2,0,0,2,52,220,439,175]
e: 453
d: [2,455,2,2,1,1,40,439,175]
d: [2,457,2,0,0,2,21,455,220]
e: 457
r: 439
i: [1,200,0,13,4,0,0,1] d2_funct_1
s: 200
a: [3,438,1,4,1,435]
s: 438
i: [1,167,0,15,4,0,0,2] t57_relat_1
s: 167
i: [1,201,0,15,4,0,0,1] d2_funct_1
s: 201
a: [3,443,1,4,1,435]
s: 443
d: [2,478,2,0,0,3,52,220,215,443,183]
e: 478
d: [2,480,2,4,2,1,40,443,183]
d: [2,482,2,2,1,2,21,480,215]
d: [2,483,2,0,0,3,21,482,220]
e: 483
r: 443
i: [1,157,0,18,5,0,0,1] d2_funct_1
s: 157
a: [3,442,1,4,1,435]
s: 442
i: [1,168,0,19,4,0,0,1] t47_funct_1
s: 168
a: [3,451,1,5,1,435]
s: 451
d: [2,499,2,11,4,1,44,451,199]
d: [2,500,2,8,3,1,19,499]
d: [2,501,2,6,2,2,21,500,215]
d: [2,502,2,4,1,3,21,501,220]
a: [3,506,2,4,1,502]
s: 506
a: [3,446,1,6,1,435]
s: 446
d: [2,526,3,0,0,4,52,220,210,506,446,167]
e: 526
d: [2,527,2,9,3,1,40,446,167]
d: [2,529,2,5,2,2,21,527,506]
d: [2,530,2,2,1,3,21,529,210]
d: [2,531,2,0,0,4,21,530,220]
e: 531
d: [2,489,1,17,4,1,44,205,168]
d: [2,493,1,10,3,1,19,489]
d: [2,494,1,8,2,2,21,493,215]
d: [2,495,1,6,1,3,21,494,220]
f: 531,502,495,483,457,435,133,134,135,136
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_1__t49_funct_1
% SZS output start Proof for funct_1__t49_funct_1
1. ! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r2_hidden(X0,X1) => k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0))) [input t49_funct_1]
2. ~! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r2_hidden(X0,X1) => k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0))) [negated conjecture 1]
12. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ! [X2] : ((~r2_hidden(X1,k9_xtuple_0(X0)) => (k1_funct_1(X0,X1) = X2 <=> k1_xboole_0 = X2)) & (r2_hidden(X1,k9_xtuple_0(X0)) => (k1_funct_1(X0,X1) = X2 <=> r2_hidden(k4_tarski(X1,X2),X0))))) [input d2_funct_1]
20. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k5_relat_1(X0,X1))) [input dt_k5_relat_1]
38. ! [X0,X1] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1)))) [input fc8_funct_1]
52. ! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) => k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0))) [input t47_funct_1]
53. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) <=> (r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1)))) [input t57_relat_1]
57. ~! [X0,X1,X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r2_hidden(X0,X1) => k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0))) [flattening 2]
58. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1,X2] : ((~r2_hidden(X1,k9_xtuple_0(X0)) => (k1_funct_1(X0,X1) = X2 <=> k1_xboole_0 = X2)) & (r2_hidden(X1,k9_xtuple_0(X0)) => (k1_funct_1(X0,X1) = X2 <=> r2_hidden(k4_tarski(X1,X2),X0))))) [flattening 12]
64. ! [X0,X1,X2] : (v1_relat_1(X2) => (r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) <=> (r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1)))) [flattening 53]
65. ! [X0,X1,X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) => k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0))) [flattening 52]
70. ? [X0,X1,X2] : ((k1_funct_1(k5_relat_1(X2,X1),X0) != k1_funct_1(X2,X0) & r2_hidden(X0,X1)) & (v1_funct_1(X2) & v1_relat_1(X2))) [ennf transformation 57]
71. ? [X0,X1,X2] : (k1_funct_1(k5_relat_1(X2,X1),X0) != k1_funct_1(X2,X0) & r2_hidden(X0,X1) & v1_funct_1(X2) & v1_relat_1(X2)) [flattening 70]
89. ! [X0] : (! [X1,X2] : (((k1_funct_1(X0,X1) = X2 <=> k1_xboole_0 = X2) | r2_hidden(X1,k9_xtuple_0(X0))) & ((k1_funct_1(X0,X1) = X2 <=> r2_hidden(k4_tarski(X1,X2),X0)) | ~r2_hidden(X1,k9_xtuple_0(X0)))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 58]
90. ! [X0] : (! [X1,X2] : (((k1_funct_1(X0,X1) = X2 <=> k1_xboole_0 = X2) | r2_hidden(X1,k9_xtuple_0(X0))) & ((k1_funct_1(X0,X1) = X2 <=> r2_hidden(k4_tarski(X1,X2),X0)) | ~r2_hidden(X1,k9_xtuple_0(X0)))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 89]
96. ! [X0,X1,X2] : ((r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) <=> (r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1))) | ~v1_relat_1(X2)) [ennf transformation 64]
97. ! [X0,X1,X2] : ((k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1)))) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) [ennf transformation 65]
98. ! [X0,X1,X2] : (k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) [flattening 97]
99. ! [X0,X1] : (v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 20]
106. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 38]
107. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 106]
110. ? [X0,X1,X2] : (k1_funct_1(k5_relat_1(X2,X1),X0) != k1_funct_1(X2,X0) & r2_hidden(X0,X1) & v1_funct_1(X2) & v1_relat_1(X2)) => (k1_funct_1(k5_relat_1(sK2,sK1),sK0) != k1_funct_1(sK2,sK0) & r2_hidden(sK0,sK1) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
111. k1_funct_1(k5_relat_1(sK2,sK1),sK0) != k1_funct_1(sK2,sK0) & r2_hidden(sK0,sK1) & v1_funct_1(sK2) & v1_relat_1(sK2) [skolemisation 71,110]
112. ! [X0] : (! [X1,X2] : ((((k1_funct_1(X0,X1) = X2 | k1_xboole_0 != X2) & (k1_xboole_0 = X2 | k1_funct_1(X0,X1) != X2)) | r2_hidden(X1,k9_xtuple_0(X0))) & (((k1_funct_1(X0,X1) = X2 | ~r2_hidden(k4_tarski(X1,X2),X0)) & (r2_hidden(k4_tarski(X1,X2),X0) | k1_funct_1(X0,X1) != X2)) | ~r2_hidden(X1,k9_xtuple_0(X0)))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 90]
113. ! [X0,X1,X2] : (((r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | (~r2_hidden(X0,k9_xtuple_0(X2)) | ~r2_hidden(X0,X1))) & ((r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1)) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))))) | ~v1_relat_1(X2)) [nnf transformation 96]
114. ! [X0,X1,X2] : (((r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | ~r2_hidden(X0,k9_xtuple_0(X2)) | ~r2_hidden(X0,X1)) & ((r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1)) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))))) | ~v1_relat_1(X2)) [flattening 113]
133. v1_relat_1(sK2) [cnf transformation 111] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
134. v1_funct_1(sK2) [cnf transformation 111] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
135. r2_hidden(sK0,sK1) [cnf transformation 111] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
136. k1_funct_1(k5_relat_1(sK2,sK1),sK0) != k1_funct_1(sK2,sK0) [cnf transformation 111] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
159. k1_funct_1(X0,X1) = X2 | k1_xboole_0 != X2 | r2_hidden(X1,k9_xtuple_0(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 112] {a:0,w:16,wCS:6400,sine:1,thAx:0,allAx:1,thDist:-1}
167. r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | ~r2_hidden(X0,k9_xtuple_0(X2)) | ~r2_hidden(X0,X1) | ~v1_relat_1(X2) [cnf transformation 114] {a:0,w:15,wCS:6000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
168. k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [cnf transformation 98] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
175. v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 99] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
183. v1_funct_1(k5_relat_1(X0,X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 107] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
199. k1_funct_1(X0,X1) = k1_xboole_0 | r2_hidden(X1,k9_xtuple_0(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [equality resolution 159] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
203. 1 <=> k1_funct_1(k5_relat_1(sK2,sK1),sK0) = k1_funct_1(sK2,sK0) [avatar definition]
205. k1_funct_1(k5_relat_1(sK2,sK1),sK0) != k1_funct_1(sK2,sK0) <- (~1) [avatar component clause 203] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
206. ~1 [avatar split clause 136,203]
208. 2 <=> r2_hidden(sK0,sK1) [avatar definition]
210. r2_hidden(sK0,sK1) <- (2) [avatar component clause 208] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
211. 2 [avatar split clause 135,208]
213. 3 <=> v1_funct_1(sK2) [avatar definition]
215. v1_funct_1(sK2) <- (3) [avatar component clause 213] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
216. 3 [avatar split clause 134,213]
218. 4 <=> v1_relat_1(sK2) [avatar definition]
220. v1_relat_1(sK2) <- (4) [avatar component clause 218] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
221. 4 [avatar split clause 133,218]
435. k1_xboole_0 != k1_funct_1(sK2,sK0) | r2_hidden(sK0,k9_xtuple_0(k5_relat_1(sK2,sK1))) | ~v1_funct_1(k5_relat_1(sK2,sK1)) | ~v1_relat_1(k5_relat_1(sK2,sK1)) <- (~1) [superposition 205,199] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
437. 31 <=> v1_relat_1(k5_relat_1(sK2,sK1)) [avatar definition]
439. ~v1_relat_1(k5_relat_1(sK2,sK1)) <- (~31) [avatar component clause 437] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
441. 32 <=> v1_funct_1(k5_relat_1(sK2,sK1)) [avatar definition]
443. ~v1_funct_1(k5_relat_1(sK2,sK1)) <- (~32) [avatar component clause 441] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
445. 33 <=> r2_hidden(sK0,k9_xtuple_0(k5_relat_1(sK2,sK1))) [avatar definition]
446. ~r2_hidden(sK0,k9_xtuple_0(k5_relat_1(sK2,sK1))) <- (~33) [avatar component clause 445] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
449. 34 <=> k1_xboole_0 = k1_funct_1(sK2,sK0) [avatar definition]
451. k1_xboole_0 != k1_funct_1(sK2,sK0) <- (~34) [avatar component clause 449] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
452. ~31 | ~32 | 33 | ~34 | 1 [avatar split clause 435,203,449,445,441,437]
455. ~v1_relat_1(sK2) <- (~31) [resolution 439,175] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
457. $false <- (4, ~31) [subsumption resolution 455,220] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
458. ~4 | 31 [avatar contradiction clause 457]
480. ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (~32) [resolution 443,183] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
482. ~v1_relat_1(sK2) <- (3, ~32) [subsumption resolution 480,215] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
483. $false <- (3, 4, ~32) [subsumption resolution 482,220] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
484. ~3 | ~4 | 32 [avatar contradiction clause 483]
489. k1_funct_1(sK2,sK0) != k1_funct_1(sK2,sK0) | ~r2_hidden(sK0,k9_xtuple_0(k5_relat_1(sK2,sK1))) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (~1) [superposition 205,168] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
493. ~r2_hidden(sK0,k9_xtuple_0(k5_relat_1(sK2,sK1))) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (~1) [trivial inequality removal 489] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
494. ~r2_hidden(sK0,k9_xtuple_0(k5_relat_1(sK2,sK1))) | ~v1_relat_1(sK2) <- (~1, 3) [subsumption resolution 493,215] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
495. ~r2_hidden(sK0,k9_xtuple_0(k5_relat_1(sK2,sK1))) <- (~1, 3, 4) [subsumption resolution 494,220] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
496. ~33 | 1 | ~3 | ~4 [avatar split clause 495,218,213,203,445]
499. k1_xboole_0 != k1_xboole_0 | r2_hidden(sK0,k9_xtuple_0(sK2)) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (~34) [superposition 451,199] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
500. r2_hidden(sK0,k9_xtuple_0(sK2)) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (~34) [trivial inequality removal 499] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
501. r2_hidden(sK0,k9_xtuple_0(sK2)) | ~v1_relat_1(sK2) <- (3, ~34) [subsumption resolution 500,215] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
502. r2_hidden(sK0,k9_xtuple_0(sK2)) <- (3, 4, ~34) [subsumption resolution 501,220] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
504. 35 <=> r2_hidden(sK0,k9_xtuple_0(sK2)) [avatar definition]
506. r2_hidden(sK0,k9_xtuple_0(sK2)) <- (35) [avatar component clause 504] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
507. 35 | ~3 | ~4 | 34 [avatar split clause 502,449,218,213,504]
527. ~r2_hidden(sK0,k9_xtuple_0(sK2)) | ~r2_hidden(sK0,sK1) | ~v1_relat_1(sK2) <- (~33) [resolution 446,167] {a:2,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
529. ~r2_hidden(sK0,sK1) | ~v1_relat_1(sK2) <- (~33, 35) [subsumption resolution 527,506] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
530. ~v1_relat_1(sK2) <- (2, ~33, 35) [subsumption resolution 529,210] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
531. $false <- (2, 4, ~33, 35) [subsumption resolution 530,220] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
532. ~2 | ~4 | 33 | ~35 [avatar contradiction clause 531]
533. $false [avatar sat refutation 206,211,216,221,452,458,484,496,507,532] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:187,thDist:-187}
% SZS output end Proof for funct_1__t49_funct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 55
% >>> Preprocessing
% Introduced skolems: 12
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 66
% Generated clauses: 243
% Active clauses: 84
% Passive clauses: 118
% Final active clauses: 81
% Final passive clauses: 25
% Main loop iterations started: 93
% 
% >>> Simplifying Inferences
% Duplicate literals: 17
% Trivial inequalities: 2
% Fw subsumption resolutions: 16
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 26
% Equational tautologies: 1
% Forward subsumptions: 32
% 
% >>> Generating Inferences
% Binary resolution: 51
% Unit resulting resolution: 3
% Forward superposition: 25
% Backward superposition: 23
% Self superposition: 3
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 4
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 44
% SAT solver unit clauses: 19
% SAT solver binary clauses: 16
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

