i: [1,165,0,2,1,1,0,0] t4_gcd_1
a: [3,254,0,2,1,165]
s: 254
i: [1,164,0,2,1,1,0,0] t4_gcd_1
a: [3,259,0,2,1,164]
s: 259
i: [1,163,0,2,1,1,0,0] t4_gcd_1
a: [3,264,0,2,1,163]
s: 264
i: [1,171,0,4,1,1,0,0] t4_gcd_1
a: [3,224,0,4,1,171]
s: 224
i: [1,170,0,4,1,1,0,0] t4_gcd_1
a: [3,229,0,4,1,170]
s: 229
i: [1,169,0,4,1,1,0,0] t4_gcd_1
a: [3,234,0,4,1,169]
s: 234
i: [1,168,0,4,1,1,0,0] t4_gcd_1
a: [3,239,0,4,1,168]
s: 239
i: [1,167,0,4,1,1,0,0] t4_gcd_1
a: [3,244,0,4,1,167]
s: 244
i: [1,166,0,4,1,1,0,0] t4_gcd_1
a: [3,249,0,4,1,166]
s: 249
i: [1,172,0,2,1,0,0,255] spc1_boole
a: [3,269,0,2,1,172]
s: 269
i: [1,173,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,274,0,2,1,173]
s: 274
i: [1,174,0,2,1,0,0,255] fc1_xboole_0
a: [3,279,0,2,1,174]
s: 279
i: [1,212,0,2,1,0,0,2] rc2_xboole_0
a: [3,289,0,2,1,212]
s: 289
i: [1,213,0,2,1,0,0,255] existence_l3_struct_0
a: [3,294,0,2,1,213]
s: 294
i: [1,214,0,2,1,0,0,255] existence_l3_algstr_0
a: [3,299,0,2,1,214]
s: 299
i: [1,215,0,2,1,0,0,1] existence_l4_algstr_0
a: [3,304,0,2,1,215]
s: 304
i: [1,216,0,2,1,0,0,2] existence_l1_struct_0
a: [3,309,0,2,1,216]
s: 309
i: [1,217,0,2,1,0,0,2] rc1_xboole_0
a: [3,314,0,2,1,217]
s: 314
i: [1,218,0,2,1,0,0,255] rc1_membered
a: [3,319,0,2,1,218]
s: 319
i: [1,219,0,2,1,0,0,255] rc2_membered
a: [3,324,0,2,1,219]
s: 324
i: [1,220,0,2,1,0,0,255] rc3_membered
a: [3,329,0,2,1,220]
s: 329
i: [1,175,0,3,1,0,0,255] d2_xboole_0
a: [3,284,0,3,1,175]
s: 284
r: 274
i: [1,176,0,4,2,0,0,255] dt_l3_struct_0
s: 176
d: [2,332,1,2,1,1,40,176,294]
a: [3,336,1,2,1,332]
s: 336
i: [1,177,0,4,2,0,0,255] dt_l3_algstr_0
s: 177
i: [1,178,0,4,2,0,0,255] dt_l4_algstr_0
s: 178
d: [2,344,1,2,1,1,40,178,254]
a: [3,349,1,2,1,344]
s: 349
i: [1,179,0,4,2,0,0,255] dt_l4_algstr_0
s: 179
d: [2,362,1,2,1,1,40,179,254]
a: [3,367,1,2,1,362]
s: 367
i: [1,209,0,4,1,0,0,1] existence_m1_subset_1
s: 209
i: [1,194,0,5,2,0,0,255] t6_boole
s: 194
d: [2,356,2,2,1,1,40,349,176]
a: [3,360,2,2,1,356]
s: 360
i: [1,208,0,5,2,0,0,255] t7_boole
s: 208
d: [2,338,1,2,1,1,40,177,299]
a: [3,342,1,2,1,338]
s: 342
i: [1,180,0,6,3,0,0,255] cc7_struct_0
s: 180
i: [1,181,0,6,3,0,0,255] cc5_struct_0
s: 181
d: [2,345,1,2,1,1,40,178,304]
a: [3,354,1,2,1,345]
s: 354
i: [1,183,0,6,3,0,0,255] cc2_struct_0
s: 183
d: [2,363,1,2,1,1,40,179,304]
a: [3,372,1,2,1,363]
s: 372
i: [1,184,0,6,3,0,0,255] cc17_struct_0
s: 184
i: [1,185,0,6,3,0,0,255] cc6_struct_0
s: 185
d: [2,381,1,4,2,1,40,181,309]
a: [3,403,1,2,1,381]
s: 403
i: [1,186,0,6,3,0,0,255] cc1_struct_0
s: 186
d: [2,379,2,4,2,1,40,181,336]
a: [3,385,2,2,1,379]
s: 385
i: [1,189,0,6,3,0,0,255] cc4_struct_0
s: 189
i: [1,205,0,6,2,0,0,255] t1_subset
s: 205
d: [2,380,2,4,2,1,40,181,342]
a: [3,394,2,2,1,380]
s: 394
i: [1,210,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 210
d: [2,409,2,2,1,1,40,354,176]
a: [3,413,2,2,1,409]
s: 413
i: [1,187,0,7,3,0,0,255] cc8_struct_0
s: 187
i: [1,190,0,7,3,0,0,255] cc11_struct_0
s: 190
d: [2,436,3,4,2,1,40,413,183]
a: [3,441,3,2,1,436]
s: 441
i: [1,191,0,7,3,0,0,255] cc11_struct_0
s: 191
i: [1,192,0,7,3,0,0,255] cc9_struct_0
s: 192
i: [1,195,0,7,3,0,0,255] fc9_struct_0
s: 195
i: [1,196,0,7,3,0,0,255] fc6_struct_0
s: 196
i: [1,201,0,7,3,0,0,1] fc2_struct_0
s: 201
i: [1,202,0,7,3,0,0,255] fc8_struct_0
s: 202
i: [1,203,0,7,3,0,0,255] fc7_struct_0
s: 203
i: [1,204,0,7,3,0,0,1] fc1_struct_0
s: 204
i: [1,207,0,7,3,0,0,255] t8_boole
s: 207
d: [2,376,1,7,3,0,44,194,194]
s: 376
i: [1,206,0,8,3,0,0,255] t2_subset
s: 206
d: [2,462,1,7,2,1,40,206,239]
a: [3,473,1,3,1,462]
s: 473
d: [2,485,3,0,0,3,52,264,360,473,201]
e: 485
d: [2,486,2,4,2,1,40,473,201]
d: [2,488,2,2,1,2,21,486,360]
d: [2,489,2,0,0,3,21,488,264]
e: 489
r: 473
i: [1,193,0,9,4,0,0,255] cc10_struct_0
s: 193
i: [1,198,0,20,6,0,0,2] d3_gcd_1
s: 198
a: [3,472,1,3,1,462]
s: 472
i: [1,199,0,20,6,0,0,2] d3_gcd_1
s: 199
a: [3,469,1,4,1,462]
s: 469
i: [1,211,0,20,6,0,0,1] symmetry_r3_gcd_1
s: 211
i: [1,200,0,24,7,0,0,2] d3_gcd_1
s: 200
d: [2,463,1,7,2,1,40,206,244]
a: [3,478,1,4,1,463]
s: 478
i: [1,197,0,30,9,0,0,1] t2_gcd_1
s: 197
d: [2,464,1,7,2,1,40,206,249]
a: [3,483,1,4,1,464]
s: 483
d: [2,495,1,16,5,1,40,198,229]
d: [2,497,1,12,4,2,21,495,239]
d: [2,498,1,8,3,3,21,497,244]
d: [2,499,1,6,2,4,21,498,254]
d: [2,500,1,4,1,5,21,499,264]
a: [3,504,1,4,1,500]
s: 504
d: [2,496,1,16,5,1,40,198,234]
d: [2,506,1,12,4,2,21,496,244]
d: [2,507,1,8,3,3,21,506,249]
d: [2,508,1,6,2,4,21,507,254]
d: [2,509,1,4,1,5,21,508,264]
a: [3,513,1,4,1,509]
s: 513
d: [2,516,1,16,5,1,40,199,229]
d: [2,518,1,12,4,2,21,516,239]
d: [2,519,1,8,3,3,21,518,244]
d: [2,520,1,6,2,4,21,519,254]
d: [2,521,1,4,1,5,21,520,264]
a: [3,525,1,4,1,521]
s: 525
d: [2,517,1,16,5,1,40,199,234]
d: [2,527,1,12,4,2,21,517,244]
d: [2,528,1,8,3,3,21,527,249]
d: [2,529,1,6,2,4,21,528,254]
d: [2,530,1,4,1,5,21,529,264]
a: [3,534,1,4,1,530]
s: 534
d: [2,544,1,16,5,1,40,211,229]
d: [2,546,1,12,4,2,21,544,239]
d: [2,547,1,8,3,3,21,546,244]
d: [2,548,1,6,2,4,21,547,254]
d: [2,549,1,4,1,5,21,548,264]
a: [3,553,1,4,1,549]
s: 553
d: [2,545,1,16,5,1,40,211,234]
d: [2,555,1,12,4,2,21,545,244]
d: [2,556,1,8,3,3,21,555,249]
d: [2,557,1,6,2,4,21,556,254]
d: [2,558,1,4,1,5,21,557,264]
a: [3,562,1,4,1,558]
s: 562
d: [2,465,1,6,2,0,40,206,209]
s: 465
d: [2,536,2,4,1,1,40,469,210]
a: [3,542,2,4,1,536]
s: 542
d: [2,452,1,9,4,0,44,187,194]
s: 452
d: [2,564,2,4,1,1,40,478,210]
a: [3,570,2,4,1,564]
s: 570
d: [2,454,1,9,4,0,44,192,194]
s: 454
d: [2,572,2,4,1,1,40,483,210]
a: [3,578,2,4,1,572]
s: 578
d: [2,580,2,26,8,1,40,504,197]
d: [2,582,2,22,7,2,21,580,239]
d: [2,583,2,18,6,3,21,582,244]
d: [2,584,2,16,5,4,21,583,254]
d: [2,585,2,14,4,5,21,584,259]
d: [2,586,2,12,3,6,21,585,264]
s: 586
d: [2,592,2,26,8,1,40,513,197]
d: [2,594,2,22,7,2,21,592,244]
d: [2,595,2,18,6,3,21,594,249]
d: [2,596,2,16,5,4,21,595,254]
d: [2,597,2,14,4,5,21,596,259]
d: [2,598,2,12,3,6,21,597,264]
s: 598
d: [2,639,3,8,2,7,40,586,513]
d: [2,641,3,4,1,8,21,639,249]
a: [3,645,3,4,1,641]
s: 645
d: [2,604,2,26,8,1,40,525,197]
d: [2,606,2,22,7,2,21,604,244]
d: [2,607,2,18,6,3,21,606,239]
d: [2,608,2,16,5,4,21,607,254]
d: [2,609,2,14,4,5,21,608,259]
d: [2,610,2,12,3,6,21,609,264]
s: 610
d: [2,611,2,26,8,1,40,534,197]
d: [2,613,2,22,7,2,21,611,249]
d: [2,614,2,18,6,3,21,613,244]
d: [2,615,2,16,5,4,21,614,254]
d: [2,616,2,14,4,5,21,615,259]
d: [2,617,2,12,3,6,21,616,264]
s: 617
d: [2,640,3,8,2,7,40,586,525]
d: [2,647,3,4,1,7,21,640,239]
a: [3,651,3,4,1,647]
s: 651
d: [2,624,2,6,2,0,40,465,210]
s: 624
d: [2,653,3,8,2,7,40,598,534]
d: [2,654,3,4,1,7,21,653,244]
a: [3,658,3,4,1,654]
s: 658
d: [2,684,3,8,2,7,40,617,513]
d: [2,686,3,4,1,7,21,684,249]
a: [3,690,3,4,1,686]
s: 690
d: [2,661,4,20,6,1,40,645,200]
d: [2,667,4,16,5,2,21,661,249]
d: [2,668,4,12,4,3,21,667,239]
d: [2,669,4,10,3,4,21,668,254]
d: [2,670,4,8,2,5,21,669,264]
a: [3,673,4,4,1,670]
s: 673
d: [2,737,5,0,0,7,52,264,254,249,239,645,224,673,200]
e: 737
d: [2,740,5,20,6,1,40,673,200]
d: [2,747,5,16,5,2,21,740,224]
d: [2,748,5,12,4,3,21,747,645]
d: [2,749,5,8,3,4,21,748,239]
d: [2,750,5,4,2,5,21,749,249]
d: [2,751,5,2,1,6,21,750,254]
d: [2,752,5,0,0,7,21,751,264]
e: 752
d: [2,685,3,8,2,7,40,617,525]
d: [2,692,3,4,1,8,21,685,239]
f: 752,692,641,530,521,509,500,163,164,165,166,167,168,169,170,171
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gcd_1__t4_gcd_1
% SZS output start Proof for gcd_1__t4_gcd_1
1. ! [X0] : ((l4_algstr_0(X0) & v3_group_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ((r3_gcd_1(X0,X2,X3) & r3_gcd_1(X0,X1,X2)) => r3_gcd_1(X0,X1,X3)))))) [input t4_gcd_1]
2. ~! [X0] : ((l4_algstr_0(X0) & v3_group_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ((r3_gcd_1(X0,X2,X3) & r3_gcd_1(X0,X1,X2)) => r3_gcd_1(X0,X1,X3)))))) [negated conjecture 1]
29. ! [X0] : ((l4_algstr_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r3_gcd_1(X0,X1,X2) <=> (r1_gcd_1(X0,X2,X1) & r1_gcd_1(X0,X1,X2)))))) [input d3_gcd_1]
58. ! [X0] : ((l4_algstr_0(X0) & v3_group_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ((r1_gcd_1(X0,X2,X3) & r1_gcd_1(X0,X1,X2)) => r1_gcd_1(X0,X1,X3)))))) [input t2_gcd_1]
84. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r3_gcd_1(X0,X1,X3) & (r3_gcd_1(X0,X2,X3) & r3_gcd_1(X0,X1,X2))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l4_algstr_0(X0) & v3_group_1(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
85. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r3_gcd_1(X0,X1,X3) & r3_gcd_1(X0,X2,X3) & r3_gcd_1(X0,X1,X2) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l4_algstr_0(X0) & v3_group_1(X0) & ~v2_struct_0(X0)) [flattening 84]
116. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_gcd_1(X0,X1,X3) | (~r1_gcd_1(X0,X2,X3) | ~r1_gcd_1(X0,X1,X2))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l4_algstr_0(X0) | ~v3_group_1(X0) | v2_struct_0(X0))) [ennf transformation 58]
117. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (r1_gcd_1(X0,X1,X3) | ~r1_gcd_1(X0,X2,X3) | ~r1_gcd_1(X0,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l4_algstr_0(X0) | ~v3_group_1(X0) | v2_struct_0(X0)) [flattening 116]
118. ! [X0] : (! [X1] : (! [X2] : ((r3_gcd_1(X0,X1,X2) <=> (r1_gcd_1(X0,X2,X1) & r1_gcd_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l4_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 29]
119. ! [X0] : (! [X1] : (! [X2] : ((r3_gcd_1(X0,X1,X2) <=> (r1_gcd_1(X0,X2,X1) & r1_gcd_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l4_algstr_0(X0) | v2_struct_0(X0)) [flattening 118]
136. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r3_gcd_1(X0,X1,X3) & r3_gcd_1(X0,X2,X3) & r3_gcd_1(X0,X1,X2) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l4_algstr_0(X0) & v3_group_1(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r3_gcd_1(sK0,X1,X3) & r3_gcd_1(sK0,X2,X3) & r3_gcd_1(sK0,X1,X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l4_algstr_0(sK0) & v3_group_1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
137. ? [X1] : (? [X2] : (? [X3] : (~r3_gcd_1(sK0,X1,X3) & r3_gcd_1(sK0,X2,X3) & r3_gcd_1(sK0,X1,X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (~r3_gcd_1(sK0,sK1,X3) & r3_gcd_1(sK0,X2,X3) & r3_gcd_1(sK0,sK1,X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
138. ? [X2] : (? [X3] : (~r3_gcd_1(sK0,sK1,X3) & r3_gcd_1(sK0,X2,X3) & r3_gcd_1(sK0,sK1,X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (~r3_gcd_1(sK0,sK1,X3) & r3_gcd_1(sK0,sK2,X3) & r3_gcd_1(sK0,sK1,sK2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
139. ? [X3] : (~r3_gcd_1(sK0,sK1,X3) & r3_gcd_1(sK0,sK2,X3) & r3_gcd_1(sK0,sK1,sK2) & m1_subset_1(X3,u1_struct_0(sK0))) => (~r3_gcd_1(sK0,sK1,sK3) & r3_gcd_1(sK0,sK2,sK3) & r3_gcd_1(sK0,sK1,sK2) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
140. (((~r3_gcd_1(sK0,sK1,sK3) & r3_gcd_1(sK0,sK2,sK3) & r3_gcd_1(sK0,sK1,sK2) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l4_algstr_0(sK0) & v3_group_1(sK0) & ~v2_struct_0(sK0) [skolemisation 85,139,138,137,136]
141. ! [X0] : (! [X1] : (! [X2] : (((r3_gcd_1(X0,X1,X2) | (~r1_gcd_1(X0,X2,X1) | ~r1_gcd_1(X0,X1,X2))) & ((r1_gcd_1(X0,X2,X1) & r1_gcd_1(X0,X1,X2)) | ~r3_gcd_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l4_algstr_0(X0) | v2_struct_0(X0)) [nnf transformation 119]
142. ! [X0] : (! [X1] : (! [X2] : (((r3_gcd_1(X0,X1,X2) | ~r1_gcd_1(X0,X2,X1) | ~r1_gcd_1(X0,X1,X2)) & ((r1_gcd_1(X0,X2,X1) & r1_gcd_1(X0,X1,X2)) | ~r3_gcd_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l4_algstr_0(X0) | v2_struct_0(X0)) [flattening 141]
163. ~v2_struct_0(sK0) [cnf transformation 140] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
164. v3_group_1(sK0) [cnf transformation 140] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
165. l4_algstr_0(sK0) [cnf transformation 140] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
166. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 140] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
167. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 140] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
168. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 140] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
169. r3_gcd_1(sK0,sK1,sK2) [cnf transformation 140] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
170. r3_gcd_1(sK0,sK2,sK3) [cnf transformation 140] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
171. ~r3_gcd_1(sK0,sK1,sK3) [cnf transformation 140] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
197. ~r1_gcd_1(X0,X2,X3) | r1_gcd_1(X0,X1,X3) | ~r1_gcd_1(X0,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l4_algstr_0(X0) | ~v3_group_1(X0) | v2_struct_0(X0) [cnf transformation 117] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
198. ~r3_gcd_1(X0,X1,X2) | r1_gcd_1(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l4_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 142] {a:0,w:20,wCS:8000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
199. ~r3_gcd_1(X0,X1,X2) | r1_gcd_1(X0,X2,X1) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l4_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 142] {a:0,w:20,wCS:8000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
200. ~r1_gcd_1(X0,X2,X1) | r3_gcd_1(X0,X1,X2) | ~r1_gcd_1(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l4_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 142] {a:0,w:24,wCS:9600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
222. 1 <=> r3_gcd_1(sK0,sK1,sK3) [avatar definition]
224. ~r3_gcd_1(sK0,sK1,sK3) <- (~1) [avatar component clause 222] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
225. ~1 [avatar split clause 171,222]
227. 2 <=> r3_gcd_1(sK0,sK2,sK3) [avatar definition]
229. r3_gcd_1(sK0,sK2,sK3) <- (2) [avatar component clause 227] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
230. 2 [avatar split clause 170,227]
232. 3 <=> r3_gcd_1(sK0,sK1,sK2) [avatar definition]
234. r3_gcd_1(sK0,sK1,sK2) <- (3) [avatar component clause 232] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
235. 3 [avatar split clause 169,232]
237. 4 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
239. m1_subset_1(sK3,u1_struct_0(sK0)) <- (4) [avatar component clause 237] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
240. 4 [avatar split clause 168,237]
242. 5 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
244. m1_subset_1(sK2,u1_struct_0(sK0)) <- (5) [avatar component clause 242] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
245. 5 [avatar split clause 167,242]
247. 6 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
249. m1_subset_1(sK1,u1_struct_0(sK0)) <- (6) [avatar component clause 247] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
250. 6 [avatar split clause 166,247]
252. 7 <=> l4_algstr_0(sK0) [avatar definition]
254. l4_algstr_0(sK0) <- (7) [avatar component clause 252] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
255. 7 [avatar split clause 165,252]
257. 8 <=> v3_group_1(sK0) [avatar definition]
259. v3_group_1(sK0) <- (8) [avatar component clause 257] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
260. 8 [avatar split clause 164,257]
262. 9 <=> v2_struct_0(sK0) [avatar definition]
264. ~v2_struct_0(sK0) <- (~9) [avatar component clause 262] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
265. ~9 [avatar split clause 163,262]
495. r1_gcd_1(sK0,sK2,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (2) [resolution 198,229] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
496. r1_gcd_1(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (3) [resolution 198,234] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
497. r1_gcd_1(sK0,sK2,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (2, 4) [subsumption resolution 495,239] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
498. r1_gcd_1(sK0,sK2,sK3) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (2, 4, 5) [subsumption resolution 497,244] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
499. r1_gcd_1(sK0,sK2,sK3) | v2_struct_0(sK0) <- (2, 4, 5, 7) [subsumption resolution 498,254] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
500. r1_gcd_1(sK0,sK2,sK3) <- (2, 4, 5, 7, ~9) [subsumption resolution 499,264] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
502. 47 <=> r1_gcd_1(sK0,sK2,sK3) [avatar definition]
504. r1_gcd_1(sK0,sK2,sK3) <- (47) [avatar component clause 502] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
505. 47 | ~2 | ~4 | ~5 | ~7 | 9 [avatar split clause 500,262,252,242,237,227,502]
506. r1_gcd_1(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (3, 5) [subsumption resolution 496,244] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
507. r1_gcd_1(sK0,sK1,sK2) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (3, 5, 6) [subsumption resolution 506,249] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
508. r1_gcd_1(sK0,sK1,sK2) | v2_struct_0(sK0) <- (3, 5, 6, 7) [subsumption resolution 507,254] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
509. r1_gcd_1(sK0,sK1,sK2) <- (3, 5, 6, 7, ~9) [subsumption resolution 508,264] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
511. 48 <=> r1_gcd_1(sK0,sK1,sK2) [avatar definition]
513. r1_gcd_1(sK0,sK1,sK2) <- (48) [avatar component clause 511] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
514. 48 | ~3 | ~5 | ~6 | ~7 | 9 [avatar split clause 509,262,252,247,242,232,511]
516. r1_gcd_1(sK0,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (2) [resolution 199,229] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
517. r1_gcd_1(sK0,sK2,sK1) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (3) [resolution 199,234] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
518. r1_gcd_1(sK0,sK3,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (2, 4) [subsumption resolution 516,239] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
519. r1_gcd_1(sK0,sK3,sK2) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (2, 4, 5) [subsumption resolution 518,244] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
520. r1_gcd_1(sK0,sK3,sK2) | v2_struct_0(sK0) <- (2, 4, 5, 7) [subsumption resolution 519,254] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
521. r1_gcd_1(sK0,sK3,sK2) <- (2, 4, 5, 7, ~9) [subsumption resolution 520,264] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
523. 49 <=> r1_gcd_1(sK0,sK3,sK2) [avatar definition]
525. r1_gcd_1(sK0,sK3,sK2) <- (49) [avatar component clause 523] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
526. 49 | ~2 | ~4 | ~5 | ~7 | 9 [avatar split clause 521,262,252,242,237,227,523]
527. r1_gcd_1(sK0,sK2,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (3, 5) [subsumption resolution 517,244] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
528. r1_gcd_1(sK0,sK2,sK1) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (3, 5, 6) [subsumption resolution 527,249] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
529. r1_gcd_1(sK0,sK2,sK1) | v2_struct_0(sK0) <- (3, 5, 6, 7) [subsumption resolution 528,254] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
530. r1_gcd_1(sK0,sK2,sK1) <- (3, 5, 6, 7, ~9) [subsumption resolution 529,264] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
532. 50 <=> r1_gcd_1(sK0,sK2,sK1) [avatar definition]
534. r1_gcd_1(sK0,sK2,sK1) <- (50) [avatar component clause 532] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
535. 50 | ~3 | ~5 | ~6 | ~7 | 9 [avatar split clause 530,262,252,247,242,232,532]
580. r1_gcd_1(sK0,X0,sK3) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | ~v3_group_1(sK0) | v2_struct_0(sK0) <- (47) [resolution 504,197] {a:2,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
582. r1_gcd_1(sK0,X0,sK3) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | ~v3_group_1(sK0) | v2_struct_0(sK0) <- (4, 47) [subsumption resolution 580,239] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
583. r1_gcd_1(sK0,X0,sK3) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | ~v3_group_1(sK0) | v2_struct_0(sK0) <- (4, 5, 47) [subsumption resolution 582,244] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
584. r1_gcd_1(sK0,X0,sK3) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v3_group_1(sK0) | v2_struct_0(sK0) <- (4, 5, 7, 47) [subsumption resolution 583,254] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
585. r1_gcd_1(sK0,X0,sK3) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (4, 5, 7, 8, 47) [subsumption resolution 584,259] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
586. ~r1_gcd_1(sK0,X0,sK2) | r1_gcd_1(sK0,X0,sK3) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (4, 5, 7, 8, ~9, 47) [subsumption resolution 585,264] {a:2,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
611. r1_gcd_1(sK0,X0,sK1) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | ~v3_group_1(sK0) | v2_struct_0(sK0) <- (50) [resolution 534,197] {a:2,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
613. r1_gcd_1(sK0,X0,sK1) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | ~v3_group_1(sK0) | v2_struct_0(sK0) <- (6, 50) [subsumption resolution 611,249] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
614. r1_gcd_1(sK0,X0,sK1) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | ~v3_group_1(sK0) | v2_struct_0(sK0) <- (5, 6, 50) [subsumption resolution 613,244] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
615. r1_gcd_1(sK0,X0,sK1) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v3_group_1(sK0) | v2_struct_0(sK0) <- (5, 6, 7, 50) [subsumption resolution 614,254] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
616. r1_gcd_1(sK0,X0,sK1) | ~r1_gcd_1(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (5, 6, 7, 8, 50) [subsumption resolution 615,259] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
617. ~r1_gcd_1(sK0,X0,sK2) | r1_gcd_1(sK0,X0,sK1) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (5, 6, 7, 8, ~9, 50) [subsumption resolution 616,264] {a:2,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
639. r1_gcd_1(sK0,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) <- (4, 5, 7, 8, ~9, 47, 48) [resolution 586,513] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
641. r1_gcd_1(sK0,sK1,sK3) <- (4, 5, 6, 7, 8, ~9, 47, 48) [subsumption resolution 639,249] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
643. 56 <=> r1_gcd_1(sK0,sK1,sK3) [avatar definition]
645. r1_gcd_1(sK0,sK1,sK3) <- (56) [avatar component clause 643] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
646. 56 | ~4 | ~5 | ~6 | ~7 | ~8 | 9 | ~47 | ~48 [avatar split clause 641,511,502,262,257,252,247,242,237,643]
672. 59 <=> r1_gcd_1(sK0,sK3,sK1) [avatar definition]
673. r1_gcd_1(sK0,sK3,sK1) <- (59) [avatar component clause 672] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
685. r1_gcd_1(sK0,sK3,sK1) | ~m1_subset_1(sK3,u1_struct_0(sK0)) <- (5, 6, 7, 8, ~9, 49, 50) [resolution 617,525] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
692. r1_gcd_1(sK0,sK3,sK1) <- (4, 5, 6, 7, 8, ~9, 49, 50) [subsumption resolution 685,239] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
693. 59 | ~4 | ~5 | ~6 | ~7 | ~8 | 9 | ~49 | ~50 [avatar split clause 692,532,523,262,257,252,247,242,237,672]
740. r3_gcd_1(sK0,sK1,sK3) | ~r1_gcd_1(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (59) [resolution 673,200] {a:5,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
747. ~r1_gcd_1(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (~1, 59) [subsumption resolution 740,224] {a:5,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
748. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (~1, 56, 59) [subsumption resolution 747,645] {a:5,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
749. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (~1, 4, 56, 59) [subsumption resolution 748,239] {a:5,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
750. ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (~1, 4, 6, 56, 59) [subsumption resolution 749,249] {a:5,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
751. v2_struct_0(sK0) <- (~1, 4, 6, 7, 56, 59) [subsumption resolution 750,254] {a:5,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
752. $false <- (~1, 4, 6, 7, ~9, 56, 59) [subsumption resolution 751,264] {a:5,w:0,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
753. 1 | ~4 | ~6 | ~7 | 9 | ~56 | ~59 [avatar contradiction clause 752]
754. $false [avatar sat refutation 225,230,235,240,245,250,255,260,265,505,514,526,535,646,693,753] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:665,thDist:-665}
% SZS output end Proof for gcd_1__t4_gcd_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 61
% >>> Preprocessing
% Introduced skolems: 14
% Pure predicates: 13
% 
% >>> Saturation
% Initial clauses: 58
% Generated clauses: 325
% Active clauses: 91
% Passive clauses: 102
% Final active clauses: 89
% Final passive clauses: 5
% Main loop iterations started: 96
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 92
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 16
% Forward subsumptions: 33
% 
% >>> Generating Inferences
% Binary resolution: 92
% Unit resulting resolution: 2
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 14
% Split components: 28
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 70
% SAT solver unit clauses: 22
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 5373
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% SAT solver time: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

