i: [1,293,0,3,1,1,0,0] t1_rat_1
a: [3,466,0,3,1,293]
s: 466
i: [1,295,0,2,1,0,0,1] fc3_numbers
a: [3,471,0,2,1,295]
s: 471
i: [1,296,0,2,1,0,0,2] fc1_numbers
a: [3,476,0,2,1,296]
s: 476
i: [1,297,0,2,1,0,0,3] spc1_boole
a: [3,481,0,2,1,297]
s: 481
i: [1,298,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,486,0,2,1,298]
s: 486
i: [1,299,0,2,1,0,0,2] fc1_xboole_0
a: [3,491,0,2,1,299]
s: 491
i: [1,300,0,2,1,0,0,255] spc0_boole
a: [3,496,0,2,1,300]
s: 496
i: [1,313,0,2,1,0,0,255] fc6_ordinal1
a: [3,556,0,2,1,313]
s: 556
i: [1,312,0,2,1,0,0,255] fc6_ordinal1
a: [3,561,0,2,1,312]
s: 561
i: [1,314,0,2,1,0,0,2] spc1_numerals
a: [3,581,0,2,1,314]
s: 581
i: [1,419,0,2,1,0,0,2] rc2_xboole_0
a: [3,601,0,2,1,419]
s: 601
i: [1,420,0,2,1,0,0,1] rc2_int_1
a: [3,606,0,2,1,420]
s: 606
i: [1,421,0,2,1,0,0,2] rc4_ordinal1
a: [3,611,0,2,1,421]
s: 611
i: [1,422,0,2,1,0,0,3] rc1_ordinal1
a: [3,616,0,2,1,422]
s: 616
i: [1,423,0,2,1,0,0,2] rc1_xreal_0
a: [3,621,0,2,1,423]
s: 621
i: [1,424,0,2,1,0,0,2] rc1_xxreal_0
a: [3,626,0,2,1,424]
s: 626
i: [1,425,0,2,1,0,0,2] rc1_xboole_0
a: [3,631,0,2,1,425]
s: 631
i: [1,429,0,2,1,0,0,255] rc2_ordinal1
a: [3,636,0,2,1,429]
s: 636
i: [1,428,0,2,1,0,0,255] rc2_ordinal1
a: [3,641,0,2,1,428]
s: 641
i: [1,427,0,2,1,0,0,255] rc2_ordinal1
a: [3,646,0,2,1,427]
s: 646
i: [1,426,0,2,1,0,0,255] rc2_ordinal1
a: [3,651,0,2,1,426]
s: 651
i: [1,431,0,2,1,0,0,2] rc5_ordinal1
a: [3,656,0,2,1,431]
s: 656
i: [1,430,0,2,1,0,0,2] rc5_ordinal1
a: [3,661,0,2,1,430]
s: 661
i: [1,435,0,2,1,0,0,2] rc3_xreal_0
a: [3,666,0,2,1,435]
s: 666
i: [1,434,0,2,1,0,0,2] rc3_xreal_0
a: [3,671,0,2,1,434]
s: 671
i: [1,433,0,2,1,0,0,2] rc3_xreal_0
a: [3,676,0,2,1,433]
s: 676
i: [1,432,0,2,1,0,0,2] rc3_xreal_0
a: [3,681,0,2,1,432]
s: 681
i: [1,439,0,2,1,0,0,2] rc2_xreal_0
a: [3,686,0,2,1,439]
s: 686
i: [1,438,0,2,1,0,0,2] rc2_xreal_0
a: [3,691,0,2,1,438]
s: 691
i: [1,437,0,2,1,0,0,2] rc2_xreal_0
a: [3,696,0,2,1,437]
s: 696
i: [1,436,0,2,1,0,0,2] rc2_xreal_0
a: [3,701,0,2,1,436]
s: 701
i: [1,441,0,2,1,0,0,2] rc3_xxreal_0
a: [3,706,0,2,1,441]
s: 706
i: [1,440,0,2,1,0,0,2] rc3_xxreal_0
a: [3,711,0,2,1,440]
s: 711
i: [1,443,0,2,1,0,0,2] rc2_xxreal_0
a: [3,716,0,2,1,443]
s: 716
i: [1,442,0,2,1,0,0,2] rc2_xxreal_0
a: [3,721,0,2,1,442]
s: 721
i: [1,447,0,2,1,0,0,2] rc4_xreal_0
a: [3,726,0,2,1,447]
s: 726
i: [1,446,0,2,1,0,0,2] rc4_xreal_0
a: [3,731,0,2,1,446]
s: 731
i: [1,445,0,2,1,0,0,2] rc4_xreal_0
a: [3,736,0,2,1,445]
s: 736
i: [1,444,0,2,1,0,0,2] rc4_xreal_0
a: [3,741,0,2,1,444]
s: 741
i: [1,449,0,2,1,0,0,2] rc4_xxreal_0
a: [3,746,0,2,1,449]
s: 746
i: [1,448,0,2,1,0,0,2] rc4_xxreal_0
a: [3,751,0,2,1,448]
s: 751
i: [1,454,0,2,1,0,0,1] rc1_int_1
a: [3,756,0,2,1,454]
s: 756
i: [1,453,0,2,1,0,0,1] rc1_int_1
a: [3,761,0,2,1,453]
s: 761
i: [1,452,0,2,1,0,0,1] rc1_int_1
a: [3,766,0,2,1,452]
s: 766
i: [1,451,0,2,1,0,0,1] rc1_int_1
a: [3,771,0,2,1,451]
s: 771
i: [1,457,0,2,1,0,0,2] rc2_nat_1
a: [3,781,0,2,1,457]
s: 781
i: [1,456,0,2,1,0,0,2] rc2_nat_1
a: [3,786,0,2,1,456]
s: 786
i: [1,294,0,12,4,1,0,0] t1_rat_1
s: 294
i: [1,393,0,3,1,0,0,255] reflexivity_r1_tarski
s: 393
i: [1,301,0,3,1,0,0,255] d2_xboole_0
a: [3,501,0,3,1,301]
s: 501
r: 486
i: [1,302,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,506,0,3,1,302]
s: 506
r: 501
r: 491
d: [2,795,0,2,1,2,23,491,506]
a: [3,804,0,2,1,795]
s: 804
i: [1,303,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,511,0,3,1,303]
s: 511
r: 561
r: 556
d: [2,806,0,2,1,2,23,561,511]
a: [3,811,0,2,1,806]
s: 811
d: [2,807,0,2,1,2,23,556,511]
a: [3,816,0,2,1,807]
s: 816
i: [1,304,0,3,1,0,0,255] t1_numerals
a: [3,516,0,3,1,304]
s: 516
i: [1,317,0,3,1,0,0,2] spc1_numerals
a: [3,566,0,3,1,317]
s: 566
i: [1,316,0,3,1,0,0,2] spc1_numerals
a: [3,571,0,3,1,316]
s: 571
i: [1,320,0,3,1,0,0,255] spc0_numerals
a: [3,586,0,3,1,320]
s: 586
i: [1,319,0,3,1,0,0,255] spc0_numerals
a: [3,591,0,3,1,319]
s: 591
i: [1,450,0,3,1,0,0,1] rc1_int_1
a: [3,776,0,3,1,450]
s: 776
d: [2,794,0,3,1,2,23,501,506]
a: [3,799,0,3,1,794]
s: 799
d: [2,818,0,3,1,2,22,516,511]
d: [2,819,0,3,1,3,22,818,506]
a: [3,823,0,3,1,819]
s: 823
i: [1,322,0,4,2,0,0,1] cc3_int_1
s: 322
i: [1,323,0,4,2,0,0,1] cc2_int_1
s: 323
i: [1,324,0,4,2,0,0,3] cc6_ordinal1
s: 324
d: [2,825,1,2,1,1,40,322,606]
a: [3,830,1,2,1,825]
s: 830
i: [1,325,0,4,2,0,0,2] cc2_xreal_0
s: 325
d: [2,832,1,2,1,1,40,323,611]
a: [3,837,1,2,1,832]
s: 837
i: [1,326,0,4,2,0,0,2] cc2_xxreal_0
s: 326
i: [1,328,0,4,2,0,0,2] cc3_nat_1
s: 328
d: [2,833,1,2,1,1,40,323,656]
a: [3,842,1,2,1,833]
s: 842
i: [1,329,0,4,2,0,0,3] cc1_nat_1
s: 329
d: [2,844,1,2,1,1,40,324,611]
a: [3,849,1,2,1,844]
s: 849
i: [1,331,0,4,2,0,0,255] cc1_ordinal1
s: 331
i: [1,332,0,4,2,0,0,255] cc1_ordinal1
s: 332
d: [2,845,1,2,1,1,40,324,656]
a: [3,854,1,2,1,845]
s: 854
i: [1,342,0,4,2,0,0,2] cc3_xreal_0
s: 342
d: [2,856,1,2,1,1,40,325,611]
a: [3,861,1,2,1,856]
s: 861
i: [1,343,0,4,2,0,0,2] cc4_xreal_0
s: 343
i: [1,346,0,4,2,0,0,255] cc9_ordinal1
s: 346
d: [2,857,1,2,1,1,40,325,656]
a: [3,866,1,2,1,857]
s: 866
i: [1,347,0,4,2,0,0,2] cc7_ordinal1
s: 347
d: [2,870,1,2,1,1,40,326,611]
a: [3,875,1,2,1,870]
s: 875
i: [1,348,0,4,2,0,0,3] cc3_ordinal1
s: 348
i: [1,392,0,4,1,0,0,2] existence_m1_subset_1
s: 392
d: [2,871,1,2,1,1,40,326,656]
a: [3,880,1,2,1,871]
s: 880
i: [1,305,0,4,1,0,0,2] dt_k5_numbers
a: [3,521,0,4,1,305]
s: 521
d: [2,882,1,2,1,1,40,328,671]
a: [3,887,1,2,1,882]
s: 887
i: [1,306,0,4,1,0,0,1] dt_k6_numbers
a: [3,526,0,4,1,306]
s: 526
i: [1,315,0,4,1,0,0,2] spc1_numerals
a: [3,576,0,4,1,315]
s: 576
d: [2,883,1,2,1,1,40,328,706]
a: [3,892,1,2,1,883]
s: 892
i: [1,318,0,4,1,0,0,255] spc0_numerals
a: [3,596,0,4,1,318]
s: 596
d: [2,896,1,2,1,1,40,331,816]
a: [3,904,1,2,1,896]
s: 904
i: [1,455,0,4,1,0,0,2] rc2_nat_1
a: [3,791,0,4,1,455]
s: 791
i: [1,334,0,5,2,0,0,4] dt_k5_xcmplx_0
s: 334
d: [2,898,1,2,1,1,40,331,616]
a: [3,914,1,2,1,898]
s: 914
i: [1,344,0,5,2,0,0,4] fc4_xreal_0
s: 344
d: [2,900,1,2,1,1,40,331,781]
a: [3,919,1,2,1,900]
s: 919
i: [1,345,0,5,2,0,0,4] fc4_xreal_0
s: 345
i: [1,349,0,5,2,0,0,2] t6_boole
s: 349
d: [2,921,1,2,1,1,40,332,816]
a: [3,929,1,2,1,921]
s: 929
i: [1,350,0,5,2,0,0,255] cc8_ordinal1
s: 350
d: [2,923,1,2,1,1,40,332,616]
a: [3,939,1,2,1,923]
s: 939
i: [1,351,0,5,2,0,0,2] cc2_nat_1
s: 351
i: [1,352,0,5,2,0,0,2] cc1_xreal_0
s: 352
d: [2,925,1,2,1,1,40,332,781]
a: [3,944,1,2,1,925]
s: 944
i: [1,389,0,5,2,0,0,1] t7_boole
s: 389
d: [2,959,1,2,1,1,40,342,621]
a: [3,972,1,2,1,959]
s: 972
i: [1,307,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r0_r0_r0
a: [3,531,0,5,1,307]
s: 531
i: [1,308,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r0_r1_r0
a: [3,536,0,5,1,308]
s: 536
d: [2,982,1,2,1,1,40,343,621]
a: [3,996,1,2,1,982]
s: 996
i: [1,309,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r0_r0
a: [3,541,0,5,1,309]
s: 541
d: [2,998,1,2,1,1,40,346,804]
a: [3,1006,1,2,1,998]
s: 1006
i: [1,310,0,5,1,0,0,3] rqRealDiv__k7_xcmplx_0__r1_r1_r1
a: [3,546,0,5,1,310]
s: 546
i: [1,311,0,5,1,0,0,3] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,551,0,5,1,311]
s: 551
d: [2,1144,1,10,4,1,44,294,546]
d: [2,1145,1,8,3,1,18,1144]
a: [3,1149,1,2,1,1145]
s: 1149
d: [2,1098,0,5,2,1,22,349,506]
s: 1098
d: [2,999,1,2,1,1,40,346,496]
a: [3,1011,1,2,1,999]
s: 1011
d: [2,1099,0,5,2,1,22,350,511]
s: 1099
i: [1,361,0,6,3,0,0,255] cc2_ordinal1
s: 361
d: [2,1000,1,2,1,1,40,346,631]
a: [3,1016,1,2,1,1000]
s: 1016
i: [1,366,0,6,3,0,0,2] cc5_xxreal_0
s: 366
d: [2,1001,1,2,1,1,40,346,741]
a: [3,1021,1,2,1,1001]
s: 1021
i: [1,368,0,6,3,0,0,2] cc5_xxreal_0
s: 368
i: [1,369,0,6,3,0,0,2] cc3_xxreal_0
s: 369
d: [2,1002,1,2,1,1,40,346,751]
a: [3,1026,1,2,1,1002]
s: 1026
i: [1,371,0,6,3,0,0,2] cc3_xxreal_0
s: 371
d: [2,1036,1,2,1,1,40,347,804]
a: [3,1044,1,2,1,1036]
s: 1044
i: [1,375,0,6,3,0,0,2] cc7_xxreal_0
s: 375
i: [1,376,0,6,3,0,0,2] cc7_xxreal_0
s: 376
d: [2,1037,1,2,1,1,40,347,496]
a: [3,1049,1,2,1,1037]
s: 1049
i: [1,385,0,6,2,0,0,1] t1_subset
s: 385
d: [2,1284,1,3,1,1,40,385,466]
a: [3,1288,1,3,1,1284]
s: 1288
i: [1,394,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 394
i: [1,461,0,6,2,0,0,1] d1_rat_1
s: 461
d: [2,1290,1,3,1,1,40,394,466]
a: [3,1294,1,3,1,1290]
s: 1294
i: [1,462,0,6,2,0,0,1] d1_rat_1
s: 462
d: [2,1296,1,3,1,1,40,461,466]
a: [3,1300,1,3,1,1296]
s: 1300
i: [1,333,0,7,3,0,0,3] cc5_ordinal1
s: 333
i: [1,335,0,7,2,0,0,1] t5_arithm
s: 335
d: [2,1320,2,4,2,1,40,333,1288]
a: [3,1343,2,2,1,1320]
s: 1343
i: [1,336,0,7,2,0,0,1] t2_arithm
s: 336
d: [2,1302,1,3,1,1,40,462,466]
a: [3,1306,1,3,1,1302]
s: 1306
i: [1,337,0,7,2,0,0,4] involutiveness_k5_xcmplx_0
s: 337
i: [1,338,0,7,2,0,0,3] t6_arithm
s: 338
d: [2,1355,1,12,5,0,44,294,335]
a: [3,1362,1,3,1,1355]
s: 1362
i: [1,339,0,7,2,0,0,3] t3_arithm
s: 339
d: [2,1308,2,3,1,1,40,1300,322]
a: [3,1312,2,3,1,1308]
s: 1312
i: [1,353,0,7,3,0,0,4] fc30_xreal_0
s: 353
i: [1,354,0,7,3,0,0,4] fc30_xreal_0
s: 354
d: [2,1365,2,3,1,1,40,1306,322]
a: [3,1369,2,3,1,1365]
s: 1369
i: [1,355,0,7,3,0,0,4] fc28_xreal_0
s: 355
d: [2,1375,3,3,1,1,40,1312,343]
a: [3,1380,3,3,1,1375]
s: 1380
i: [1,356,0,7,3,0,0,4] fc28_xreal_0
s: 356
i: [1,362,0,7,3,0,0,4] fc29_xreal_0
s: 362
d: [2,1376,3,3,1,1,40,1312,342]
a: [3,1385,3,3,1,1376]
s: 1385
i: [1,363,0,7,3,0,0,4] fc29_xreal_0
s: 363
d: [2,1388,3,3,1,1,40,1369,343]
a: [3,1393,3,3,1,1388]
s: 1393
i: [1,372,0,7,3,0,0,4] fc27_xreal_0
s: 372
i: [1,373,0,7,3,0,0,4] fc27_xreal_0
s: 373
d: [2,1389,3,3,1,1,40,1369,342]
a: [3,1398,3,3,1,1389]
s: 1398
i: [1,387,0,7,2,0,0,255] t3_subset
s: 387
d: [2,1038,1,2,1,1,40,347,631]
a: [3,1054,1,2,1,1038]
s: 1054
i: [1,388,0,7,3,0,0,1] t8_boole
s: 388
i: [1,321,0,8,3,0,0,255] cc10_ordinal1
s: 321
d: [2,1039,1,2,1,1,40,347,741]
a: [3,1059,1,2,1,1039]
s: 1059
i: [1,340,0,8,2,0,0,3] spc3_arithm
s: 340
d: [2,1040,1,2,1,1,40,347,751]
a: [3,1064,1,2,1,1040]
s: 1064
i: [1,358,0,8,4,0,0,2] cc4_xxreal_0
s: 358
i: [1,360,0,8,4,0,0,2] cc6_xxreal_0
s: 360
d: [2,1489,4,9,3,1,40,358,1393]
a: [3,1505,4,3,1,1489]
s: 1505
r: 1505
i: [1,364,0,8,4,0,0,2] cc8_xxreal_0
s: 364
a: [3,1504,4,3,1,1489]
s: 1504
i: [1,386,0,8,3,0,0,1] t2_subset
s: 386
i: [1,407,0,8,3,0,0,1] fc2_int_1
s: 407
a: [3,1509,4,3,1,1489]
s: 1509
i: [1,408,0,8,3,0,0,2] fc2_nat_1
s: 408
a: [3,1512,4,3,1,1489]
s: 1512
i: [1,413,0,8,3,0,0,1] fc8_xreal_0
s: 413
i: [1,414,0,8,3,0,0,2] fc6_xreal_0
s: 414
d: [2,1490,4,9,3,1,40,358,1380]
a: [3,1518,4,3,1,1490]
s: 1518
r: 1518
i: [1,391,0,9,3,0,0,1] t5_subset
s: 391
a: [3,1517,4,3,1,1490]
s: 1517
i: [1,459,0,9,3,0,0,1] d1_rat_1
s: 459
i: [1,381,0,10,3,0,0,1] d1_rat_1
s: 381
a: [3,1522,4,3,1,1490]
s: 1522
i: [1,382,0,10,3,0,0,1] d1_rat_1
s: 382
a: [3,1525,4,3,1,1490]
s: 1525
i: [1,390,0,10,3,0,0,1] t4_subset
s: 390
i: [1,415,0,10,3,0,0,255] dt_k10_real_1
s: 415
d: [2,1066,1,2,1,1,40,348,804]
a: [3,1074,1,2,1,1066]
s: 1074
i: [1,460,0,10,2,0,0,1] d1_rat_1
s: 460
d: [2,1067,1,2,1,1,40,348,496]
a: [3,1079,1,2,1,1067]
s: 1079
i: [1,409,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 409
i: [1,341,0,12,3,0,0,4] d9_xcmplx_0
s: 341
d: [2,1068,1,2,1,1,40,348,631]
a: [3,1084,1,2,1,1068]
s: 1084
i: [1,395,0,12,5,0,0,1] fc33_xreal_0
s: 395
d: [2,1069,1,2,1,1,40,348,741]
a: [3,1089,1,2,1,1069]
s: 1089
i: [1,396,0,12,5,0,0,2] fc26_xreal_0
s: 396
i: [1,397,0,12,5,0,0,1] fc31_xreal_0
s: 397
d: [2,1070,1,2,1,1,40,348,751]
a: [3,1094,1,2,1,1070]
s: 1094
i: [1,398,0,12,5,0,0,1] fc32_xreal_0
s: 398
d: [2,1100,1,2,1,1,40,351,823]
a: [3,1107,1,2,1,1100]
s: 1107
i: [1,399,0,12,5,0,0,2] fc23_xreal_0
s: 399
i: [1,400,0,12,5,0,0,2] fc24_xreal_0
s: 400
d: [2,1101,1,2,1,1,40,351,571]
a: [3,1112,1,2,1,1101]
s: 1112
i: [1,401,0,12,5,0,0,1] fc34_xreal_0
s: 401
d: [2,1102,1,2,1,1,40,351,591]
a: [3,1117,1,2,1,1102]
s: 1117
i: [1,402,0,12,5,0,0,2] fc25_xreal_0
s: 402
i: [1,410,0,12,3,0,0,4] spc12_arithm
s: 410
d: [2,1124,1,2,1,1,40,352,566]
a: [3,1131,1,2,1,1124]
s: 1131
i: [1,416,0,12,3,0,0,255] redefinition_k10_real_1
s: 416
d: [2,1125,1,2,1,1,40,352,586]
a: [3,1136,1,2,1,1125]
s: 1136
i: [1,411,0,13,3,0,0,4] spc11_arithm
s: 411
i: [1,406,0,14,4,0,0,2] existence_m2_subset_1
s: 406
d: [2,1165,1,2,1,2,40,1099,571]
a: [3,1172,1,2,1,1165]
s: 1172
d: [2,2027,2,0,0,2,52,1149,1172,323]
e: 2027
d: [2,2031,2,2,1,1,40,1172,323]
d: [2,2033,2,0,0,2,21,2031,1149]
e: 2033
r: 1149
i: [1,412,0,14,3,0,0,4] spc10_arithm
s: 412
a: [3,1148,1,2,1,1145]
s: 1148
i: [1,383,0,15,3,0,0,1] d1_rat_1
s: 383
i: [1,403,0,15,5,0,0,2] dt_m2_subset_1
s: 403
a: [3,1152,1,3,1,1145]
s: 1152
i: [1,404,0,15,5,0,0,2] redefinition_m2_subset_1
s: 404
a: [3,1157,1,3,1,1145]
s: 1157
i: [1,405,0,15,5,0,0,2] redefinition_m2_subset_1
s: 405
i: [1,384,0,17,5,0,0,1] d1_rat_1
s: 384
d: [2,1374,1,12,5,0,44,294,338]
a: [3,2055,1,7,3,1374]
s: 2055
i: [1,417,0,17,4,0,0,1] spc4_arithm
s: 417
d: [2,2146,2,4,2,1,47,2055]
a: [3,2150,2,2,1,2146]
s: 2150
i: [1,418,0,17,4,0,0,2] spc7_arithm
s: 418
d: [2,1184,1,4,2,1,40,366,804]
a: [3,1191,1,2,1,1184]
s: 1191
d: [2,1185,1,4,2,1,40,366,496]
a: [3,1196,1,2,1,1185]
s: 1196
d: [2,1186,1,4,2,1,40,366,631]
a: [3,1201,1,2,1,1186]
s: 1201
a: [3,1206,1,2,1,1186]
s: 1206
d: [2,1187,1,4,2,1,40,366,741]
d: [2,1208,1,2,1,2,21,1187,731]
a: [3,1212,1,2,1,1208]
s: 1212
d: [2,1188,1,4,2,1,40,366,751]
d: [2,1214,1,2,1,2,21,1188,746]
a: [3,1218,1,2,1,1214]
s: 1218
d: [2,1220,1,4,2,1,40,368,581]
a: [3,1225,1,2,1,1220]
s: 1225
d: [2,1221,1,4,2,1,40,368,691]
d: [2,1228,1,2,1,2,21,1221,696]
a: [3,1232,1,2,1,1228]
s: 1232
d: [2,1222,1,4,2,1,40,368,716]
d: [2,1234,1,2,1,2,21,1222,721]
a: [3,1238,1,2,1,1234]
s: 1238
d: [2,1241,1,4,2,1,40,369,691]
d: [2,1243,1,2,1,2,21,1241,696]
a: [3,1247,1,2,1,1243]
s: 1247
d: [2,1242,1,4,2,1,40,369,716]
d: [2,1249,1,2,1,2,21,1242,721]
a: [3,1253,1,2,1,1249]
s: 1253
d: [2,1316,1,4,2,1,40,333,566]
a: [3,1329,1,2,1,1316]
s: 1329
a: [3,1333,1,2,1,1316]
s: 1333
d: [2,1428,1,4,2,1,40,321,521]
a: [3,1434,1,2,1,1428]
s: 1434
d: [2,1455,1,6,2,1,44,340,546]
a: [3,1464,1,2,1,1455]
s: 1464
d: [2,1493,1,6,3,1,40,358,626]
a: [3,1557,1,2,1,1493]
s: 1557
r: 1557
a: [3,1556,1,2,1,1493]
s: 1556
a: [3,1561,1,2,1,1493]
s: 1561
a: [3,1564,1,2,1,1493]
s: 1564
d: [2,1501,1,6,3,1,40,358,771]
a: [3,1580,1,2,1,1501]
s: 1580
r: 1580
a: [3,1579,1,2,1,1501]
s: 1579
a: [3,1584,1,2,1,1501]
s: 1584
a: [3,1587,1,2,1,1501]
s: 1587
d: [2,1815,1,11,4,0,44,409,336]
d: [2,1837,1,9,3,0,18,1815]
a: [3,1846,1,2,1,1837]
s: 1846
d: [2,1920,1,12,6,0,44,397,335]
d: [2,1925,1,10,5,0,21,1920,342]
a: [3,1932,1,2,1,1925]
s: 1932
d: [2,1798,1,16,5,0,44,294,460]
d: [2,1799,1,13,4,0,21,1798,461]
d: [2,1800,1,10,3,0,21,1799,462]
s: 1800
r: 1517
r: 1522
r: 1846
s: 1518
a: [3,1845,1,2,1,1837]
s: 1845
d: [2,1715,5,3,1,1,40,1518,348]
a: [3,1723,5,3,1,1715]
s: 1723
d: [2,1458,1,13,5,0,44,294,340]
a: [3,2051,1,11,4,1458]
s: 2051
d: [2,2130,1,15,5,1,44,384,546]
d: [2,2134,1,13,4,1,18,2130]
d: [2,2140,1,11,3,2,21,2134,1148]
s: 2140
d: [2,1716,5,3,1,1,40,1518,347]
a: [3,1728,5,3,1,1716]
s: 1728
d: [2,1103,1,3,1,0,40,351,392]
a: [3,1122,1,3,1,1103]
s: 1122
d: [2,1717,5,3,1,1,40,1518,346]
a: [3,1733,5,3,1,1717]
s: 1733
d: [2,1126,1,3,1,0,40,352,392]
a: [3,1141,1,3,1,1126]
s: 1141
d: [2,1167,1,3,1,1,40,1099,392]
a: [3,1178,1,3,1,1167]
s: 1178
d: [2,2376,6,3,1,1,40,1723,332]
a: [3,2381,6,3,1,2376]
s: 2381
d: [2,1612,1,5,2,1,40,386,823]
d: [2,1622,1,3,1,2,21,1612,811]
a: [3,1626,1,3,1,1622]
s: 1626
d: [2,2377,6,3,1,1,40,1723,331]
a: [3,2386,6,3,1,2377]
s: 2386
d: [2,1613,1,7,2,1,40,386,521]
a: [3,1635,1,3,1,1613]
s: 1635
r: 1635
a: [3,1634,1,3,1,1613]
s: 1634
d: [2,2358,2,7,2,0,47,1800]
d: [2,2359,2,4,1,1,21,2358,466]
a: [3,2363,2,4,1,2359]
s: 2363
r: 1300
r: 1312
r: 1380
r: 1385
r: 1518
r: 1525
r: 1723
r: 1728
r: 1733
r: 2381
r: 2386
d: [2,1614,1,5,2,1,40,386,566]
d: [2,1637,1,3,1,2,21,1614,476]
a: [3,1641,1,3,1,1637]
s: 1641
d: [2,2479,3,9,2,1,44,460,2363]
d: [2,2484,3,6,1,2,21,2479,466]
a: [3,2488,3,6,1,2484]
s: 2488
d: [2,1615,1,5,2,1,40,386,571]
d: [2,1643,1,3,1,2,21,1615,811]
a: [3,1647,1,3,1,1643]
s: 1647
d: [2,1616,1,5,2,1,40,386,586]
d: [2,1649,1,3,1,2,21,1616,476]
a: [3,1653,1,3,1,1649]
s: 1653
d: [2,2501,4,12,5,1,44,395,2488]
d: [2,2513,4,10,4,2,21,2501,1107]
d: [2,2514,4,7,3,3,21,2513,1369]
d: [2,2515,4,4,2,4,21,2514,1512]
a: [3,2519,4,2,1,2515]
s: 2519
d: [2,1617,1,5,2,1,40,386,591]
d: [2,1655,1,3,1,2,21,1617,811]
a: [3,1659,1,3,1,1655]
s: 1659
d: [2,2502,4,12,5,1,44,397,2488]
d: [2,2521,4,10,4,2,21,2502,1932]
d: [2,2522,4,7,3,3,21,2521,1369]
d: [2,2523,4,4,2,4,21,2522,1512]
a: [3,2527,4,2,1,2523]
s: 2527
d: [2,1620,1,5,2,1,40,386,776]
d: [2,1661,1,3,1,2,21,1620,476]
a: [3,1665,1,3,1,1661]
s: 1665
d: [2,1744,1,7,3,1,44,459,546]
d: [2,1748,1,5,2,1,18,1744]
a: [3,2047,1,3,1,1748]
s: 2047
d: [2,2505,4,7,3,1,44,413,2488]
d: [2,2529,4,4,2,2,21,2505,1369]
a: [3,2533,4,2,1,2529]
s: 2533
d: [2,2101,1,11,4,1,40,403,526]
d: [2,2106,1,7,3,2,21,2101,521]
d: [2,2107,1,5,2,3,21,2106,811]
d: [2,2108,1,3,1,4,21,2107,476]
a: [3,2112,1,3,1,2108]
s: 2112
a: [3,2153,2,2,1,2146]
s: 2153
d: [2,2090,1,21,6,0,44,294,383]
d: [2,2097,1,18,5,0,21,2090,381]
d: [2,2098,1,15,4,0,21,2097,382]
s: 2098
d: [2,2127,1,17,6,0,44,384,338]
d: [2,2135,1,15,5,1,21,2127,1148]
s: 2135
d: [2,2594,5,2,1,1,40,2533,343]
a: [3,2599,5,2,1,2594]
s: 2599
r: 1509
r: 1512
r: 2519
r: 1932
a: [3,1931,1,2,1,1925]
s: 1931
d: [2,2630,2,0,0,3,52,1191,804,1931,369]
e: 2630
d: [2,2631,2,4,2,1,40,1931,369]
d: [2,2635,2,2,1,2,21,2631,804]
d: [2,2636,2,0,0,3,21,2635,1191]
e: 2636
r: 1931
s: 2519
s: 1932
d: [2,2131,1,18,6,0,44,384,340]
d: [2,2141,1,16,5,1,21,2131,1148]
s: 2141
d: [2,2614,6,6,3,1,40,2599,358]
d: [2,2615,6,4,2,2,21,2614,2519]
d: [2,2616,6,2,1,3,21,2615,2527]
a: [3,2620,6,2,1,2616]
s: 2620
d: [2,2643,7,0,0,3,52,804,1362,2620,388]
e: 2643
d: [2,2644,7,0,0,3,52,804,1362,2620,388]
e: 2644
d: [2,2645,7,0,0,3,52,1362,2620,1098]
e: 2645
r: 1362
a: [3,1358,1,7,3,1355]
s: 1358
r: 2051
d: [2,2664,2,0,0,4,52,1148,1464,1152,1358]
e: 2664
r: 1358
r: 2527
r: 2620
d: [2,1164,1,2,1,2,40,1099,823]
d: [2,1255,2,2,1,1,40,1044,326]
d: [2,1256,2,2,1,1,40,1044,325]
d: [2,1258,2,2,1,1,40,1044,323]
d: [2,1682,5,6,2,1,40,1509,368]
d: [2,1683,5,3,1,2,21,1682,1393]
d: [2,2003,2,2,1,1,40,1131,342]
d: [2,2504,4,12,5,1,44,401,2488]
d: [2,2628,4,9,4,2,21,2504,1369]
d: [2,2684,4,10,4,2,21,2502,1932]
d: [2,2685,4,7,3,3,21,2684,1369]
d: [2,2687,6,4,2,2,21,2614,2519]
d: [2,2503,4,12,5,1,44,398,2488]
d: [2,2689,4,9,4,2,21,2503,1369]
d: [2,2690,4,7,3,3,21,2689,1107]
f: 2690,2687,2685,2664,2645,2636,2628,2514,2594,2529,2484,2359,2031,2003,1683,1388,1365,1355,1302,1258,1256,1255,1165,1164,1124,1100,819,795,316,317,304,303,302,301,299,298,297,293
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rat_1__t1_rat_1
% SZS output start Proof for rat_1__t1_rat_1
1. ! [X0] : ~(! [X1] : (v1_int_1(X1) => ! [X2] : (v1_int_1(X2) => ~(k7_xcmplx_0(X1,X2) = X0 & k6_numbers != X2))) & r2_hidden(X0,k3_numbers)) [input t1_rat_1]
2. ~! [X0] : ~(! [X1] : (v1_int_1(X1) => ! [X2] : (v1_int_1(X2) => ~(k7_xcmplx_0(X1,X2) = X0 & k6_numbers != X2))) & r2_hidden(X0,k3_numbers)) [negated conjecture 1]
7. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xreal_0(X0)) [input cc1_xreal_0]
8. ! [X0] : (v7_ordinal1(X0) => v1_int_1(X0)) [input cc2_int_1]
9. ! [X0] : (m1_subset_1(X0,k5_numbers) => ~v3_xxreal_0(X0)) [input cc2_nat_1]
11. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
12. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
13. ! [X0] : (v1_int_1(X0) => v1_xreal_0(X0)) [input cc3_int_1]
16. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
17. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
19. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
20. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) => (v2_xxreal_0(X0) & v1_xxreal_0(X0))) [input cc4_xxreal_0]
22. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
27. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
31. ! [X0] : (k3_numbers = X0 <=> ! [X1] : (r2_hidden(X1,X0) <=> ? [X2] : (? [X3] : (k7_xcmplx_0(X2,X3) = X1 & v1_int_1(X3)) & v1_int_1(X2)))) [input d1_rat_1]
32. k1_xboole_0 = o_0_0_xboole_0 [input d2_xboole_0]
47. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
51. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
62. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v2_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc31_xreal_0]
63. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v2_xxreal_0(k7_xcmplx_0(X1,X0))) [input fc32_xreal_0]
64. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc33_xreal_0]
65. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc34_xreal_0]
72. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
92. k5_numbers = k4_ordinal1 [input redefinition_k5_numbers]
93. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
106. ~v1_xboole_0(1) [input spc1_boole]
107. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
111. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
118. ! [X0] : (v1_xcmplx_0(X0) => k6_numbers = k7_xcmplx_0(k6_numbers,X0)) [input t5_arithm]
121. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
136. ? [X0] : (! [X1] : (! [X2] : ((k7_xcmplx_0(X1,X2) != X0 | k6_numbers = X2) | ~v1_int_1(X2)) | ~v1_int_1(X1)) & r2_hidden(X0,k3_numbers)) [ennf transformation 2]
137. ? [X0] : (! [X1] : (! [X2] : (k7_xcmplx_0(X1,X2) != X0 | k6_numbers = X2 | ~v1_int_1(X2)) | ~v1_int_1(X1)) & r2_hidden(X0,k3_numbers)) [flattening 136]
139. ! [X0] : (v1_xreal_0(X0) | ~v1_int_1(X0)) [ennf transformation 13]
140. ! [X0] : (v1_int_1(X0) | ~v7_ordinal1(X0)) [ennf transformation 8]
142. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
143. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
149. ! [X0] : (k6_numbers = k7_xcmplx_0(k6_numbers,X0) | ~v1_xcmplx_0(X0)) [ennf transformation 118]
156. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
157. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 19]
162. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 121]
163. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 27]
164. ! [X0] : (~v3_xxreal_0(X0) | ~m1_subset_1(X0,k5_numbers)) [ennf transformation 9]
165. ! [X0] : (v1_xreal_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 7]
170. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | (v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0))) [ennf transformation 20]
171. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0)) [flattening 170]
180. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 22]
181. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 180]
182. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 17]
183. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 182]
198. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 64]
199. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 198]
202. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 62]
203. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 202]
204. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X1,X0)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 63]
205. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 204]
210. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 65]
211. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 210]
232. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 72]
233. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 232]
244. ? [X0] : (! [X1] : (! [X2] : (k7_xcmplx_0(X1,X2) != X0 | k6_numbers = X2 | ~v1_int_1(X2)) | ~v1_int_1(X1)) & r2_hidden(X0,k3_numbers)) => (! [X1] : (! [X2] : (k7_xcmplx_0(X1,X2) != sK0 | k6_numbers = X2 | ~v1_int_1(X2)) | ~v1_int_1(X1)) & r2_hidden(sK0,k3_numbers)) [choice axiom]
245. ! [X1] : (! [X2] : (k7_xcmplx_0(X1,X2) != sK0 | k6_numbers = X2 | ~v1_int_1(X2)) | ~v1_int_1(X1)) & r2_hidden(sK0,k3_numbers) [skolemisation 137,244]
246. ! [X0] : ((k3_numbers = X0 | ? [X1] : ((! [X2] : (! [X3] : (k7_xcmplx_0(X2,X3) != X1 | ~v1_int_1(X3)) | ~v1_int_1(X2)) | ~r2_hidden(X1,X0)) & (? [X2] : (? [X3] : (k7_xcmplx_0(X2,X3) = X1 & v1_int_1(X3)) & v1_int_1(X2)) | r2_hidden(X1,X0)))) & (! [X1] : ((r2_hidden(X1,X0) | ! [X2] : (! [X3] : (k7_xcmplx_0(X2,X3) != X1 | ~v1_int_1(X3)) | ~v1_int_1(X2))) & (? [X2] : (? [X3] : (k7_xcmplx_0(X2,X3) = X1 & v1_int_1(X3)) & v1_int_1(X2)) | ~r2_hidden(X1,X0))) | k3_numbers != X0)) [nnf transformation 31]
247. ! [X0] : ((k3_numbers = X0 | ? [X1] : ((! [X2] : (! [X3] : (k7_xcmplx_0(X2,X3) != X1 | ~v1_int_1(X3)) | ~v1_int_1(X2)) | ~r2_hidden(X1,X0)) & (? [X4] : (? [X5] : (k7_xcmplx_0(X4,X5) = X1 & v1_int_1(X5)) & v1_int_1(X4)) | r2_hidden(X1,X0)))) & (! [X6] : ((r2_hidden(X6,X0) | ! [X7] : (! [X8] : (k7_xcmplx_0(X7,X8) != X6 | ~v1_int_1(X8)) | ~v1_int_1(X7))) & (? [X9] : (? [X10] : (k7_xcmplx_0(X9,X10) = X6 & v1_int_1(X10)) & v1_int_1(X9)) | ~r2_hidden(X6,X0))) | k3_numbers != X0)) [rectify 246]
248. ! [X0] : (? [X1] : ((! [X2] : (! [X3] : (k7_xcmplx_0(X2,X3) != X1 | ~v1_int_1(X3)) | ~v1_int_1(X2)) | ~r2_hidden(X1,X0)) & (? [X4] : (? [X5] : (k7_xcmplx_0(X4,X5) = X1 & v1_int_1(X5)) & v1_int_1(X4)) | r2_hidden(X1,X0))) => ((! [X2] : (! [X3] : (k7_xcmplx_0(X2,X3) != sK1(X0) | ~v1_int_1(X3)) | ~v1_int_1(X2)) | ~r2_hidden(sK1(X0),X0)) & (? [X4] : (? [X5] : (k7_xcmplx_0(X4,X5) = sK1(X0) & v1_int_1(X5)) & v1_int_1(X4)) | r2_hidden(sK1(X0),X0)))) [choice axiom]
249. ! [X0] : (? [X4] : (? [X5] : (k7_xcmplx_0(X4,X5) = sK1(X0) & v1_int_1(X5)) & v1_int_1(X4)) => (? [X5] : (sK1(X0) = k7_xcmplx_0(sK2(X0),X5) & v1_int_1(X5)) & v1_int_1(sK2(X0)))) [choice axiom]
250. ! [X0] : (? [X5] : (sK1(X0) = k7_xcmplx_0(sK2(X0),X5) & v1_int_1(X5)) => (sK1(X0) = k7_xcmplx_0(sK2(X0),sK3(X0)) & v1_int_1(sK3(X0)))) [choice axiom]
251. ! [X6] : (? [X9] : (? [X10] : (k7_xcmplx_0(X9,X10) = X6 & v1_int_1(X10)) & v1_int_1(X9)) => (? [X10] : (k7_xcmplx_0(sK4(X6),X10) = X6 & v1_int_1(X10)) & v1_int_1(sK4(X6)))) [choice axiom]
252. ! [X6] : (? [X10] : (k7_xcmplx_0(sK4(X6),X10) = X6 & v1_int_1(X10)) => (k7_xcmplx_0(sK4(X6),sK5(X6)) = X6 & v1_int_1(sK5(X6)))) [choice axiom]
253. ! [X0] : ((k3_numbers = X0 | ((! [X2] : (! [X3] : (k7_xcmplx_0(X2,X3) != sK1(X0) | ~v1_int_1(X3)) | ~v1_int_1(X2)) | ~r2_hidden(sK1(X0),X0)) & (((sK1(X0) = k7_xcmplx_0(sK2(X0),sK3(X0)) & v1_int_1(sK3(X0))) & v1_int_1(sK2(X0))) | r2_hidden(sK1(X0),X0)))) & (! [X6] : ((r2_hidden(X6,X0) | ! [X7] : (! [X8] : (k7_xcmplx_0(X7,X8) != X6 | ~v1_int_1(X8)) | ~v1_int_1(X7))) & (((k7_xcmplx_0(sK4(X6),sK5(X6)) = X6 & v1_int_1(sK5(X6))) & v1_int_1(sK4(X6))) | ~r2_hidden(X6,X0))) | k3_numbers != X0)) [skolemisation 247,252,251,250,249,248]
293. r2_hidden(sK0,k3_numbers) [cnf transformation 245] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
294. k7_xcmplx_0(X1,X2) != sK0 | k6_numbers = X2 | ~v1_int_1(X2) | ~v1_int_1(X1) [cnf transformation 245] {a:0,w:12,wCS:1200,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
297. ~v1_xboole_0(1) [cnf transformation 106] {a:0,w:2,wCS:800,sine:3,thAx:0,allAx:1,thDist:-1}
298. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 47] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
299. v1_xboole_0(k1_xboole_0) [cnf transformation 51] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
301. k1_xboole_0 = o_0_0_xboole_0 [cnf transformation 32] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
302. k6_numbers = k1_xboole_0 [cnf transformation 93] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
303. k5_numbers = k4_ordinal1 [cnf transformation 92] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
304. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 111] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
316. m1_subset_1(1,k5_numbers) [cnf transformation 107] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
317. m1_subset_1(1,k1_numbers) [cnf transformation 107] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
322. ~v1_int_1(X0) | v1_xreal_0(X0) [cnf transformation 139] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
323. ~v7_ordinal1(X0) | v1_int_1(X0) [cnf transformation 140] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
325. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 142] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
326. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 143] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
335. k6_numbers = k7_xcmplx_0(k6_numbers,X0) | ~v1_xcmplx_0(X0) [cnf transformation 149] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
342. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 156] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
343. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 157] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
349. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 162] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
350. v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1) [cnf transformation 163] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
351. ~m1_subset_1(X0,k5_numbers) | ~v3_xxreal_0(X0) [cnf transformation 164] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
352. ~m1_subset_1(X0,k1_numbers) | v1_xreal_0(X0) [cnf transformation 165] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
358. ~v1_xxreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xboole_0(X0) [cnf transformation 171] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
368. ~v2_xxreal_0(X0) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 181] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
369. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 183] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
377. v1_int_1(sK4(X6)) | ~r2_hidden(X6,X0) | k3_numbers != X0 [cnf transformation 253] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:6,thDist:-6}
378. v1_int_1(sK5(X6)) | ~r2_hidden(X6,X0) | k3_numbers != X0 [cnf transformation 253] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:6,thDist:-6}
379. k7_xcmplx_0(sK4(X6),sK5(X6)) = X6 | ~r2_hidden(X6,X0) | k3_numbers != X0 [cnf transformation 253] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:6,thDist:-6}
395. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 199] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
397. ~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 203] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
398. ~v2_xxreal_0(k7_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 205] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
401. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 211] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
413. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 233] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
460. k7_xcmplx_0(sK4(X6),sK5(X6)) = X6 | ~r2_hidden(X6,k3_numbers) [equality resolution 379] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
461. ~r2_hidden(X6,k3_numbers) | v1_int_1(sK5(X6)) [equality resolution 378] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
462. ~r2_hidden(X6,k3_numbers) | v1_int_1(sK4(X6)) [equality resolution 377] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
464. 1 <=> r2_hidden(sK0,k3_numbers) [avatar definition]
466. r2_hidden(sK0,k3_numbers) <- (1) [avatar component clause 464] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
467. 1 [avatar split clause 293,464]
479. 4 <=> v1_xboole_0(1) [avatar definition]
482. ~4 [avatar split clause 297,479]
484. 5 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
487. 5 [avatar split clause 298,484]
489. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
491. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 489] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
492. 6 [avatar split clause 299,489]
499. 8 <=> k1_xboole_0 = o_0_0_xboole_0 [avatar definition]
502. 8 [avatar split clause 301,499]
504. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
506. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 504] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
507. 9 [avatar split clause 302,504]
509. 10 <=> k5_numbers = k4_ordinal1 [avatar definition]
511. k5_numbers = k4_ordinal1 <- (10) [avatar component clause 509] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
512. 10 [avatar split clause 303,509]
514. 11 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
516. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (11) [avatar component clause 514] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
517. 11 [avatar split clause 304,514]
564. 21 <=> m1_subset_1(1,k1_numbers) [avatar definition]
566. m1_subset_1(1,k1_numbers) <- (21) [avatar component clause 564] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
567. 21 [avatar split clause 317,564]
569. 22 <=> m1_subset_1(1,k5_numbers) [avatar definition]
571. m1_subset_1(1,k5_numbers) <- (22) [avatar component clause 569] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
572. 22 [avatar split clause 316,569]
795. v1_xboole_0(k6_numbers) <- (6, 9) [backward demodulation 491,506] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
802. 68 <=> v1_xboole_0(k6_numbers) [avatar definition]
804. v1_xboole_0(k6_numbers) <- (68) [avatar component clause 802] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
805. 68 | ~6 | ~9 [avatar split clause 795,504,489,802]
818. m1_subset_1(k1_xboole_0,k5_numbers) <- (10, 11) [forward demodulation 516,511] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
819. m1_subset_1(k6_numbers,k5_numbers) <- (9, 10, 11) [forward demodulation 818,506] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
821. 71 <=> m1_subset_1(k6_numbers,k5_numbers) [avatar definition]
823. m1_subset_1(k6_numbers,k5_numbers) <- (71) [avatar component clause 821] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
824. 71 | ~9 | ~10 | ~11 [avatar split clause 819,514,509,504,821]
1042. 104 <=> v7_ordinal1(k6_numbers) [avatar definition]
1044. v7_ordinal1(k6_numbers) <- (104) [avatar component clause 1042] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1098. k6_numbers = X0 | ~v1_xboole_0(X0) <- (9) [forward demodulation 349,506] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1099. ~m1_subset_1(X0,k5_numbers) | v7_ordinal1(X0) <- (10) [forward demodulation 350,511] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1100. ~v3_xxreal_0(k6_numbers) <- (71) [resolution 351,823] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
1105. 114 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1107. ~v3_xxreal_0(k6_numbers) <- (~114) [avatar component clause 1105] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1108. ~114 | ~71 [avatar split clause 1100,821,1105]
1124. v1_xreal_0(1) <- (21) [resolution 352,566] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
1129. 118 <=> v1_xreal_0(1) [avatar definition]
1131. v1_xreal_0(1) <- (118) [avatar component clause 1129] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1132. 118 | ~21 [avatar split clause 1124,564,1129]
1147. 121 <=> v1_int_1(1) [avatar definition]
1148. v1_int_1(1) <- (121) [avatar component clause 1147] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1151. 122 <=> k6_numbers = 1 [avatar definition]
1152. k6_numbers != 1 <- (~122) [avatar component clause 1151] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1164. v7_ordinal1(k6_numbers) <- (10, 71) [resolution 1099,823] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1165. v7_ordinal1(1) <- (10, 22) [resolution 1099,571] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1168. 104 | ~10 | ~71 [avatar split clause 1164,821,509,1042]
1170. 124 <=> v7_ordinal1(1) [avatar definition]
1172. v7_ordinal1(1) <- (124) [avatar component clause 1170] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1173. 124 | ~10 | ~22 [avatar split clause 1165,569,509,1170]
1190. 126 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1191. v1_xxreal_0(k6_numbers) <- (126) [avatar component clause 1190] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1255. v1_xxreal_0(k6_numbers) <- (104) [resolution 1044,326] {a:2,w:2,wCS:800,sine:2,thAx:0,allAx:3,thDist:-3}
1256. v1_xreal_0(k6_numbers) <- (104) [resolution 1044,325] {a:2,w:2,wCS:800,sine:2,thAx:0,allAx:3,thDist:-3}
1258. v1_int_1(k6_numbers) <- (104) [resolution 1044,323] {a:2,w:2,wCS:800,sine:1,thAx:0,allAx:3,thDist:-3}
1260. 126 | ~104 [avatar split clause 1255,1042,1190]
1262. 137 <=> v1_xreal_0(k6_numbers) [avatar definition]
1265. 137 | ~104 [avatar split clause 1256,1042,1262]
1268. 138 <=> v1_int_1(k6_numbers) [avatar definition]
1271. 138 | ~104 [avatar split clause 1258,1042,1268]
1302. v1_int_1(sK4(sK0)) <- (1) [resolution 462,466] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1304. 143 <=> v1_int_1(sK4(sK0)) [avatar definition]
1306. v1_int_1(sK4(sK0)) <- (143) [avatar component clause 1304] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1307. 143 | ~1 [avatar split clause 1302,464,1304]
1355. k6_numbers != sK0 | k6_numbers = X0 | ~v1_int_1(X0) | ~v1_int_1(k6_numbers) | ~v1_xcmplx_0(X0) [superposition 294,335] {a:1,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
1357. 150 <=> ! [X0] : (k6_numbers = X0 | ~v1_xcmplx_0(X0) | ~v1_int_1(X0)) [avatar definition]
1358. k6_numbers = X0 | ~v1_xcmplx_0(X0) | ~v1_int_1(X0) <- (150) [avatar component clause 1357] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1360. 151 <=> k6_numbers = sK0 [avatar definition]
1362. k6_numbers != sK0 <- (~151) [avatar component clause 1360] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1363. ~138 | 150 | ~151 [avatar split clause 1355,1360,1357,1268]
1365. v1_xreal_0(sK4(sK0)) <- (143) [resolution 1306,322] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1367. 152 <=> v1_xreal_0(sK4(sK0)) [avatar definition]
1369. v1_xreal_0(sK4(sK0)) <- (152) [avatar component clause 1367] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1370. 152 | ~143 [avatar split clause 1365,1304,1367]
1388. v1_xxreal_0(sK4(sK0)) <- (152) [resolution 1369,343] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
1391. 155 <=> v1_xxreal_0(sK4(sK0)) [avatar definition]
1393. v1_xxreal_0(sK4(sK0)) <- (155) [avatar component clause 1391] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
1394. 155 | ~152 [avatar split clause 1388,1367,1391]
1463. 163 <=> v1_xcmplx_0(1) [avatar definition]
1464. v1_xcmplx_0(1) <- (163) [avatar component clause 1463] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1507. 168 <=> v2_xxreal_0(sK4(sK0)) [avatar definition]
1509. v2_xxreal_0(sK4(sK0)) <- (168) [avatar component clause 1507] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
1511. 169 <=> v3_xxreal_0(sK4(sK0)) [avatar definition]
1682. ~v3_xxreal_0(sK4(sK0)) | ~v1_xxreal_0(sK4(sK0)) <- (168) [resolution 1509,368] {a:5,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
1683. ~v3_xxreal_0(sK4(sK0)) <- (155, 168) [subsumption resolution 1682,1393] {a:5,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
1684. ~169 | ~155 | ~168 [avatar split clause 1683,1507,1391,1511]
1798. sK0 != X2 | k6_numbers = sK5(X2) | ~v1_int_1(sK5(X2)) | ~v1_int_1(sK4(X2)) | ~r2_hidden(X2,k3_numbers) [superposition 294,460] {a:1,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
1799. sK0 != X2 | k6_numbers = sK5(X2) | ~v1_int_1(sK4(X2)) | ~r2_hidden(X2,k3_numbers) [subsumption resolution 1798,461] {a:1,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
1800. sK0 != X2 | k6_numbers = sK5(X2) | ~r2_hidden(X2,k3_numbers) [subsumption resolution 1799,462] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
1930. 217 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1931. v2_xxreal_0(k6_numbers) <- (217) [avatar component clause 1930] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1932. ~v2_xxreal_0(k6_numbers) <- (~217) [avatar component clause 1930] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2003. v1_xcmplx_0(1) <- (118) [resolution 1131,342] {a:2,w:2,wCS:800,sine:2,thAx:0,allAx:3,thDist:-3}
2005. 163 | ~118 [avatar split clause 2003,1129,1463]
2031. v1_int_1(1) <- (124) [resolution 1172,323] {a:2,w:2,wCS:800,sine:1,thAx:0,allAx:3,thDist:-3}
2035. 121 | ~124 [avatar split clause 2031,1170,1147]
2059. k1_xboole_0 != o_0_0_xboole_0 | k6_numbers != k1_xboole_0 | k6_numbers != 1 | v1_xboole_0(1) | ~v1_xboole_0(o_0_0_xboole_0) [theory tautology sat conflict] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:1,thDist:-1}
2358. k6_numbers = sK5(sK0) | ~r2_hidden(sK0,k3_numbers) [equality resolution 1800] {a:2,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
2359. k6_numbers = sK5(sK0) <- (1) [subsumption resolution 2358,466] {a:2,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
2361. 249 <=> k6_numbers = sK5(sK0) [avatar definition]
2363. k6_numbers = sK5(sK0) <- (249) [avatar component clause 2361] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
2364. 249 | ~1 [avatar split clause 2359,464,2361]
2479. sK0 = k7_xcmplx_0(sK4(sK0),k6_numbers) | ~r2_hidden(sK0,k3_numbers) <- (249) [superposition 460,2363] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
2484. sK0 = k7_xcmplx_0(sK4(sK0),k6_numbers) <- (1, 249) [subsumption resolution 2479,466] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
2486. 264 <=> sK0 = k7_xcmplx_0(sK4(sK0),k6_numbers) [avatar definition]
2488. sK0 = k7_xcmplx_0(sK4(sK0),k6_numbers) <- (264) [avatar component clause 2486] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2489. 264 | ~1 | ~249 [avatar split clause 2484,2361,464,2486]
2501. ~v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK4(sK0)) | v3_xxreal_0(sK4(sK0)) <- (264) [superposition 395,2488] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2502. ~v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(sK4(sK0)) | v3_xxreal_0(sK4(sK0)) <- (264) [superposition 397,2488] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2503. ~v2_xxreal_0(sK0) | ~v1_xreal_0(sK4(sK0)) | v2_xxreal_0(sK4(sK0)) | ~v1_xreal_0(k6_numbers) | v3_xxreal_0(k6_numbers) <- (264) [superposition 398,2488] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2504. ~v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(sK4(sK0)) | v2_xxreal_0(sK4(sK0)) <- (264) [superposition 401,2488] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2505. v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK4(sK0)) <- (264) [superposition 413,2488] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2513. ~v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK4(sK0)) | v3_xxreal_0(sK4(sK0)) <- (~114, 264) [subsumption resolution 2501,1107] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2514. ~v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | v3_xxreal_0(sK4(sK0)) <- (~114, 152, 264) [subsumption resolution 2513,1369] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2517. 267 <=> v3_xxreal_0(sK0) [avatar definition]
2519. ~v3_xxreal_0(sK0) <- (~267) [avatar component clause 2517] {a:4,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:15,thDist:-15}
2525. 268 <=> v2_xxreal_0(sK0) [avatar definition]
2529. v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (152, 264) [subsumption resolution 2505,1369] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2531. 269 <=> v1_xreal_0(sK0) [avatar definition]
2533. v1_xreal_0(sK0) <- (269) [avatar component clause 2531] {a:4,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:15,thDist:-15}
2534. ~137 | 269 | ~152 | ~264 [avatar split clause 2529,2486,1367,2531,1262]
2594. v1_xxreal_0(sK0) <- (269) [resolution 2533,343] {a:5,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:16,thDist:-16}
2597. 278 <=> v1_xxreal_0(sK0) [avatar definition]
2599. v1_xxreal_0(sK0) <- (278) [avatar component clause 2597] {a:5,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:16,thDist:-16}
2600. 278 | ~269 [avatar split clause 2594,2531,2597]
2614. v3_xxreal_0(sK0) | v2_xxreal_0(sK0) | v1_xboole_0(sK0) <- (278) [resolution 2599,358] {a:6,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:17,thDist:-17}
2618. 280 <=> v1_xboole_0(sK0) [avatar definition]
2620. v1_xboole_0(sK0) <- (280) [avatar component clause 2618] {a:6,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:17,thDist:-17}
2625. 169 | ~137 | ~267 | 114 | ~152 | ~264 [avatar split clause 2514,2486,1367,1105,2517,1262,1511]
2628. ~v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | v2_xxreal_0(k6_numbers) | v2_xxreal_0(sK4(sK0)) <- (152, 264) [subsumption resolution 2504,1369] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2629. 168 | 217 | ~137 | ~267 | ~152 | ~264 [avatar split clause 2628,2486,1367,2517,1262,1930,1507]
2631. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (217) [resolution 1931,369] {a:2,w:4,wCS:1600,sine:2,thAx:0,allAx:3,thDist:-3}
2635. ~v1_xxreal_0(k6_numbers) <- (68, 217) [subsumption resolution 2631,804] {a:2,w:2,wCS:800,sine:2,thAx:0,allAx:3,thDist:-3}
2636. $false <- (68, 126, 217) [subsumption resolution 2635,1191] {a:2,w:0,sine:2,thAx:0,allAx:3,thDist:-3}
2637. ~68 | ~126 | ~217 [avatar contradiction clause 2636]
2645. $false <- (9, ~151, 280) [unit resulting resolution 1362,2620,1098] {a:7,w:0,goal:1,sine:2,thAx:0,allAx:21,thDist:-21}
2652. ~9 | 151 | ~280 [avatar contradiction clause 2645]
2664. $false <- (121, ~122, 150, 163) [unit resulting resolution 1148,1464,1152,1358] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
2683. ~121 | 122 | ~150 | ~163 [avatar contradiction clause 2664]
2684. ~v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK4(sK0)) | v3_xxreal_0(sK4(sK0)) <- (~217, 264) [subsumption resolution 2502,1932] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2685. ~v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | v3_xxreal_0(sK4(sK0)) <- (152, ~217, 264) [subsumption resolution 2684,1369] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2686. 169 | ~137 | ~268 | ~152 | 217 | ~264 [avatar split clause 2685,2486,1930,1367,2525,1262,1511]
2687. v2_xxreal_0(sK0) | v1_xboole_0(sK0) <- (~267, 278) [subsumption resolution 2614,2519] {a:6,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:17,thDist:-17}
2688. 280 | 268 | 267 | ~278 [avatar split clause 2687,2597,2517,2525,2618]
2689. ~v2_xxreal_0(sK0) | v2_xxreal_0(sK4(sK0)) | ~v1_xreal_0(k6_numbers) | v3_xxreal_0(k6_numbers) <- (152, 264) [subsumption resolution 2503,1369] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2690. ~v2_xxreal_0(sK0) | v2_xxreal_0(sK4(sK0)) | ~v1_xreal_0(k6_numbers) <- (~114, 152, 264) [subsumption resolution 2689,1107] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
2691. ~137 | 168 | ~268 | 114 | ~152 | ~264 [avatar split clause 2690,2486,1367,1105,2525,1507,1262]
2692. $false [avatar sat refutation 467,482,487,492,502,507,512,517,567,572,805,824,1108,1132,1168,1173,1260,1265,1271,1307,1363,1370,1394,1684,2005,2035,2059,2364,2489,2534,2600,2625,2629,2637,2652,2683,2686,2688,2691] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:2172,thDist:-2172}
% SZS output end Proof for rat_1__t1_rat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 122
% >>> Preprocessing
% Introduced skolems: 25
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 165
% Generated clauses: 1262
% Active clauses: 279
% Passive clauses: 534
% Final active clauses: 244
% Final passive clauses: 232
% Main loop iterations started: 293
% 
% >>> Simplifying Inferences
% Duplicate literals: 77
% Fw subsumption resolutions: 122
% Fw demodulations: 5
% Bw demodulations: 16
% 
% >>> Deletion Inferences
% Simple tautologies: 57
% Equational tautologies: 4
% Forward subsumptions: 140
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 313
% Unit resulting resolution: 6
% Forward superposition: 155
% Backward superposition: 106
% Self superposition: 3
% Equality resolution: 3
% 
% >>> AVATAR
% Split clauses: 70
% Split components: 162
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 333
% SAT solver unit clauses: 68
% SAT solver binary clauses: 166
% 
% Memory used [KB]: 6780
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% congruence closure: 0.002 s
% unit clause index maintenance: 0.001 s
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.012 s ( own 0.006 s ) 

