i: [1,23,0,10,1,1,0,0] t24_xtuple_0
a: [3,33,0,10,1,23]
s: 33
i: [1,27,0,3,1,0,0,1] reflexivity_r1_tarski
s: 27
i: [1,24,0,5,1,0,0,1] t17_xboole_1
s: 24
i: [1,28,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 28
i: [1,29,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 29
i: [1,25,0,8,2,0,0,1] t8_xtuple_0
s: 25
d: [2,36,1,5,1,0,44,24,29]
s: 36
i: [1,26,0,11,3,0,0,1] t19_xboole_1
s: 26
d: [2,41,1,14,2,1,40,26,33]
a: [3,49,1,7,1,41]
s: 49
d: [2,55,2,0,0,1,52,24,49,25]
e: 55
d: [2,56,2,5,1,1,40,49,25]
d: [2,58,2,0,0,1,21,56,24]
e: 58
r: 49
a: [3,48,1,7,1,41]
s: 48
a: [3,53,1,7,1,41]
s: 53
d: [2,60,2,0,0,1,52,36,53,25]
e: 60
d: [2,61,2,5,1,1,40,53,25]
d: [2,63,2,0,0,1,21,61,36]
e: 63
f: 63,58,41,23
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t24_xtuple_0
% SZS output start Proof for xtuple_0__t24_xtuple_0
1. ! [X0] : ! [X1] : r1_tarski(k9_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) [input t24_xtuple_0]
2. ~! [X0] : ! [X1] : r1_tarski(k9_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) [negated conjecture 1]
3. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
8. ! [X0] : ! [X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [input t17_xboole_1]
9. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k3_xboole_0(X1,X2))) [input t19_xboole_1]
10. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) [input t8_xtuple_0]
11. ~! [X0,X1] : r1_tarski(k9_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) [flattening 2]
12. ! [X0,X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [flattening 8]
13. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) [flattening 10]
14. ! [X0,X1,X2] : ((r1_tarski(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k3_xboole_0(X1,X2))) [flattening 9]
17. ? [X0,X1] : ~r1_tarski(k9_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) [ennf transformation 11]
18. ! [X0,X1] : (r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 13]
19. ! [X0,X1,X2] : (r1_tarski(X0,k3_xboole_0(X1,X2)) | (~r1_tarski(X0,X2) | ~r1_tarski(X0,X1))) [ennf transformation 14]
20. ! [X0,X1,X2] : (r1_tarski(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X2) | ~r1_tarski(X0,X1)) [flattening 19]
21. ? [X0,X1] : ~r1_tarski(k9_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) => ~r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))) [choice axiom]
22. ~r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))) [skolemisation 17,21]
23. ~r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))) [cnf transformation 22] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
24. r1_tarski(k3_xboole_0(X0,X1),X0) [cnf transformation 12] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
25. r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 18] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
26. r1_tarski(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 20] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
29. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
31. 1 <=> r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))) [avatar definition]
33. ~r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1))) <- (~1) [avatar component clause 31] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
34. ~1 [avatar split clause 23,31]
36. r1_tarski(k3_xboole_0(X1,X0),X0) [superposition 24,29] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
41. ~r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k9_xtuple_0(sK1)) | ~r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k9_xtuple_0(sK0)) <- (~1) [resolution 26,33] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
47. 2 <=> r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k9_xtuple_0(sK0)) [avatar definition]
49. ~r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k9_xtuple_0(sK0)) <- (~2) [avatar component clause 47] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
51. 3 <=> r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k9_xtuple_0(sK1)) [avatar definition]
53. ~r1_tarski(k9_xtuple_0(k3_xboole_0(sK0,sK1)),k9_xtuple_0(sK1)) <- (~3) [avatar component clause 51] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
54. ~2 | ~3 | 1 [avatar split clause 41,31,51,47]
56. ~r1_tarski(k3_xboole_0(sK0,sK1),sK0) <- (~2) [resolution 49,25] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
58. $false <- (~2) [subsumption resolution 56,24] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
59. 2 [avatar contradiction clause 58]
61. ~r1_tarski(k3_xboole_0(sK0,sK1),sK1) <- (~3) [resolution 53,25] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
63. $false <- (~3) [subsumption resolution 61,36] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
64. 3 [avatar contradiction clause 63]
65. $false [avatar sat refutation 34,54,59,64] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:28,thDist:-28}
% SZS output end Proof for xtuple_0__t24_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 7
% Generated clauses: 28
% Active clauses: 11
% Passive clauses: 12
% Final active clauses: 10
% Final passive clauses: 1
% Main loop iterations started: 11
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 2
% Forward superposition: 6
% Backward superposition: 3
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 8
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

