i: [1,79,0,3,1,1,0,0] t26_subset_1
a: [3,117,0,3,1,79]
s: 117
i: [1,78,0,3,1,1,0,0] t26_subset_1
a: [3,122,0,3,1,78]
s: 122
i: [1,77,0,4,1,1,0,0] t26_subset_1
a: [3,127,0,4,1,77]
s: 127
i: [1,76,0,4,1,1,0,0] t26_subset_1
a: [3,132,0,4,1,76]
s: 132
i: [1,75,0,4,1,1,0,0] t26_subset_1
a: [3,137,0,4,1,75]
s: 137
i: [1,80,0,5,1,1,0,0] t26_subset_1
a: [3,112,0,5,1,80]
s: 112
i: [1,81,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,142,0,2,1,81]
s: 142
i: [1,82,0,2,1,0,0,255] fc1_xboole_0
a: [3,147,0,2,1,82]
s: 147
i: [1,104,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,157,0,2,1,104]
s: 157
i: [1,105,0,2,1,0,0,2] rc2_xboole_0
a: [3,162,0,2,1,105]
s: 162
i: [1,106,0,2,1,0,0,2] rc1_xboole_0
a: [3,167,0,2,1,106]
s: 167
i: [1,108,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,172,0,2,1,108]
s: 172
i: [1,107,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,177,0,2,1,107]
s: 177
i: [1,84,0,3,1,0,0,1] fc1_subset_1
s: 84
i: [1,99,0,3,1,0,0,1] rc2_subset_1
s: 99
i: [1,100,0,3,1,0,0,1] reflexivity_r1_tarski
s: 100
i: [1,83,0,3,1,0,0,255] d2_xboole_0
a: [3,152,0,3,1,83]
s: 152
r: 142
i: [1,87,0,4,2,0,0,255] cc2_zfmisc_1
s: 87
d: [2,180,1,2,1,1,40,87,147]
a: [3,186,1,2,1,180]
s: 186
i: [1,90,0,4,2,0,0,255] cc1_zfmisc_1
s: 90
i: [1,97,0,4,1,0,0,1] existence_m1_subset_1
s: 97
d: [2,182,1,2,1,1,40,87,167]
a: [3,191,1,2,1,182]
s: 191
i: [1,85,0,5,1,0,0,255] t4_boole
s: 85
d: [2,181,1,3,1,0,40,87,99]
s: 181
i: [1,86,0,5,1,0,0,255] t3_boole
s: 86
i: [1,89,0,5,2,0,0,1] rc1_subset_1
s: 89
i: [1,91,0,5,2,0,0,255] t6_boole
s: 91
i: [1,98,0,5,1,0,0,1] rc2_subset_1
s: 98
i: [1,101,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 101
i: [1,88,0,7,2,0,0,1] rc1_subset_1
s: 88
i: [1,95,0,7,3,0,0,2] t8_boole
s: 95
d: [2,198,1,3,1,1,40,101,117]
a: [3,202,1,3,1,198]
s: 202
i: [1,96,0,9,3,0,0,1] t63_xboole_1
s: 96
d: [2,205,1,6,2,1,40,96,117]
s: 205
i: [1,102,0,10,2,0,0,1] dt_k3_subset_1
s: 102
i: [1,92,0,11,2,0,0,255] d4_subset_1
s: 92
d: [2,206,2,6,2,1,40,96,202]
s: 206
i: [1,103,0,11,2,0,0,1] involutiveness_k3_subset_1
s: 103
d: [2,212,3,3,1,2,40,206,122]
a: [3,217,3,3,1,212]
s: 217
i: [1,93,0,16,4,0,0,1] t23_subset_1
s: 93
d: [2,227,4,0,0,4,52,217,137,127,112,93]
e: 227
d: [2,228,1,11,3,1,40,93,112]
d: [2,231,1,8,2,2,21,228,217]
d: [2,232,1,4,1,3,21,231,127]
d: [2,233,1,0,0,4,21,232,137]
e: 233
f: 233,212,198,75,77,78,79,80
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t26_subset_1
% SZS output start Proof for subset_1__t26_subset_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(X0)) => ((r1_xboole_0(X3,X2) & r1_tarski(X1,X2)) => r1_tarski(X1,k3_subset_1(X0,X3)))))) [input t26_subset_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(X0)) => ((r1_xboole_0(X3,X2) & r1_tarski(X1,X2)) => r1_tarski(X1,k3_subset_1(X0,X3)))))) [negated conjecture 1]
25. ! [X0,X1] : (r1_xboole_0(X0,X1) => r1_xboole_0(X1,X0)) [input symmetry_r1_xboole_0]
26. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => (r1_xboole_0(X1,X2) <=> r1_tarski(X1,k3_subset_1(X0,X2))))) [input t23_subset_1]
29. ! [X0] : ! [X1] : ! [X2] : ((r1_xboole_0(X1,X2) & r1_tarski(X0,X1)) => r1_xboole_0(X0,X2)) [input t63_xboole_1]
33. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(X0)) => ((r1_xboole_0(X3,X2) & r1_tarski(X1,X2)) => r1_tarski(X1,k3_subset_1(X0,X3)))))) [flattening 2]
35. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => (r1_xboole_0(X1,X2) <=> r1_tarski(X1,k3_subset_1(X0,X2))))) [flattening 26]
38. ! [X0,X1,X2] : ((r1_xboole_0(X1,X2) & r1_tarski(X0,X1)) => r1_xboole_0(X0,X2)) [flattening 29]
42. ? [X0,X1] : (? [X2] : (? [X3] : ((~r1_tarski(X1,k3_subset_1(X0,X3)) & (r1_xboole_0(X3,X2) & r1_tarski(X1,X2))) & m1_subset_1(X3,k1_zfmisc_1(X0))) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 33]
43. ? [X0,X1] : (? [X2] : (? [X3] : (~r1_tarski(X1,k3_subset_1(X0,X3)) & r1_xboole_0(X3,X2) & r1_tarski(X1,X2) & m1_subset_1(X3,k1_zfmisc_1(X0))) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 42]
49. ! [X0,X1] : (! [X2] : ((r1_xboole_0(X1,X2) <=> r1_tarski(X1,k3_subset_1(X0,X2))) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 35]
51. ! [X0,X1,X2] : (r1_xboole_0(X0,X2) | (~r1_xboole_0(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 38]
52. ! [X0,X1,X2] : (r1_xboole_0(X0,X2) | ~r1_xboole_0(X1,X2) | ~r1_tarski(X0,X1)) [flattening 51]
53. ! [X0,X1] : (r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 25]
56. ? [X0,X1] : (? [X2] : (? [X3] : (~r1_tarski(X1,k3_subset_1(X0,X3)) & r1_xboole_0(X3,X2) & r1_tarski(X1,X2) & m1_subset_1(X3,k1_zfmisc_1(X0))) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (? [X2] : (? [X3] : (~r1_tarski(sK1,k3_subset_1(sK0,X3)) & r1_xboole_0(X3,X2) & r1_tarski(sK1,X2) & m1_subset_1(X3,k1_zfmisc_1(sK0))) & m1_subset_1(X2,k1_zfmisc_1(sK0))) & m1_subset_1(sK1,k1_zfmisc_1(sK0))) [choice axiom]
57. ? [X2] : (? [X3] : (~r1_tarski(sK1,k3_subset_1(sK0,X3)) & r1_xboole_0(X3,X2) & r1_tarski(sK1,X2) & m1_subset_1(X3,k1_zfmisc_1(sK0))) & m1_subset_1(X2,k1_zfmisc_1(sK0))) => (? [X3] : (~r1_tarski(sK1,k3_subset_1(sK0,X3)) & r1_xboole_0(X3,sK2) & r1_tarski(sK1,sK2) & m1_subset_1(X3,k1_zfmisc_1(sK0))) & m1_subset_1(sK2,k1_zfmisc_1(sK0))) [choice axiom]
58. ? [X3] : (~r1_tarski(sK1,k3_subset_1(sK0,X3)) & r1_xboole_0(X3,sK2) & r1_tarski(sK1,sK2) & m1_subset_1(X3,k1_zfmisc_1(sK0))) => (~r1_tarski(sK1,k3_subset_1(sK0,sK3)) & r1_xboole_0(sK3,sK2) & r1_tarski(sK1,sK2) & m1_subset_1(sK3,k1_zfmisc_1(sK0))) [choice axiom]
59. ((~r1_tarski(sK1,k3_subset_1(sK0,sK3)) & r1_xboole_0(sK3,sK2) & r1_tarski(sK1,sK2) & m1_subset_1(sK3,k1_zfmisc_1(sK0))) & m1_subset_1(sK2,k1_zfmisc_1(sK0))) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) [skolemisation 43,58,57,56]
62. ! [X0,X1] : (! [X2] : (((r1_xboole_0(X1,X2) | ~r1_tarski(X1,k3_subset_1(X0,X2))) & (r1_tarski(X1,k3_subset_1(X0,X2)) | ~r1_xboole_0(X1,X2))) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [nnf transformation 49]
75. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 59] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
77. m1_subset_1(sK3,k1_zfmisc_1(sK0)) [cnf transformation 59] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
78. r1_tarski(sK1,sK2) [cnf transformation 59] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
79. r1_xboole_0(sK3,sK2) [cnf transformation 59] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
80. ~r1_tarski(sK1,k3_subset_1(sK0,sK3)) [cnf transformation 59] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
93. r1_tarski(X1,k3_subset_1(X0,X2)) | ~r1_xboole_0(X1,X2) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 62] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
96. ~r1_xboole_0(X1,X2) | r1_xboole_0(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 52] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. ~r1_xboole_0(X0,X1) | r1_xboole_0(X1,X0) [cnf transformation 53] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
110. 1 <=> r1_tarski(sK1,k3_subset_1(sK0,sK3)) [avatar definition]
112. ~r1_tarski(sK1,k3_subset_1(sK0,sK3)) <- (~1) [avatar component clause 110] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
113. ~1 [avatar split clause 80,110]
115. 2 <=> r1_xboole_0(sK3,sK2) [avatar definition]
117. r1_xboole_0(sK3,sK2) <- (2) [avatar component clause 115] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
118. 2 [avatar split clause 79,115]
120. 3 <=> r1_tarski(sK1,sK2) [avatar definition]
122. r1_tarski(sK1,sK2) <- (3) [avatar component clause 120] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
123. 3 [avatar split clause 78,120]
125. 4 <=> m1_subset_1(sK3,k1_zfmisc_1(sK0)) [avatar definition]
127. m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (4) [avatar component clause 125] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
128. 4 [avatar split clause 77,125]
135. 6 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
137. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (6) [avatar component clause 135] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
138. 6 [avatar split clause 75,135]
198. r1_xboole_0(sK2,sK3) <- (2) [resolution 101,117] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
200. 17 <=> r1_xboole_0(sK2,sK3) [avatar definition]
202. r1_xboole_0(sK2,sK3) <- (17) [avatar component clause 200] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
203. 17 | ~2 [avatar split clause 198,115,200]
206. ~r1_tarski(X1,sK2) | r1_xboole_0(X1,sK3) <- (17) [resolution 96,202] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
212. r1_xboole_0(sK1,sK3) <- (3, 17) [resolution 206,122] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
215. 18 <=> r1_xboole_0(sK1,sK3) [avatar definition]
217. r1_xboole_0(sK1,sK3) <- (18) [avatar component clause 215] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
218. 18 | ~3 | ~17 [avatar split clause 212,200,120,215]
228. ~r1_xboole_0(sK1,sK3) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (~1) [resolution 93,112] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
231. ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (~1, 18) [subsumption resolution 228,217] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
232. ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (~1, 4, 18) [subsumption resolution 231,127] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
233. $false <- (~1, 4, 6, 18) [subsumption resolution 232,137] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
234. 1 | ~4 | ~6 | ~18 [avatar contradiction clause 233]
236. $false [avatar sat refutation 113,118,123,128,138,203,218,234] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:121,thDist:-121}
% SZS output end Proof for subset_1__t26_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 31
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 34
% Generated clauses: 83
% Active clauses: 39
% Passive clauses: 50
% Final active clauses: 38
% Final passive clauses: 10
% Main loop iterations started: 40
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 1
% Forward superposition: 3
% Backward superposition: 7
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 14
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit clause index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

