i: [1,100,0,2,1,1,0,0] t11_isomichi
a: [3,137,0,2,1,100]
s: 137
i: [1,99,0,2,1,1,0,0] t11_isomichi
a: [3,142,0,2,1,99]
s: 142
i: [1,101,0,8,2,1,0,0] t11_isomichi
a: [3,128,0,4,1,101]
s: 128
i: [1,123,0,2,1,0,0,255] existence_l1_struct_0
a: [3,147,0,2,1,123]
s: 147
i: [1,124,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,152,0,2,1,124]
s: 152
i: [1,118,0,3,1,0,0,255] reflexivity_r1_tarski
s: 118
i: [1,104,0,4,2,0,0,255] dt_l1_pre_topc
s: 104
i: [1,117,0,4,1,0,0,2] existence_m1_subset_1
s: 117
d: [2,154,1,2,1,1,40,104,137]
a: [3,159,1,2,1,154]
s: 159
i: [1,102,0,7,2,0,0,255] d3_struct_0
s: 102
i: [1,116,0,7,2,0,0,255] t3_subset
s: 116
d: [2,155,1,2,1,1,40,104,152]
a: [3,164,1,2,1,155]
s: 164
i: [1,103,0,8,2,0,0,255] dt_k2_struct_0
s: 103
i: [1,110,0,8,2,0,0,255] rc4_tops_1
s: 110
i: [1,105,0,9,2,0,0,2] t2_tops_1
s: 105
i: [1,112,0,10,3,0,0,255] rc6_pre_topc
s: 112
d: [2,166,1,10,3,0,44,103,102]
d: [2,167,1,8,2,0,18,166]
s: 167
i: [1,113,0,10,3,0,0,255] rc10_pre_topc
s: 113
d: [2,168,1,10,3,0,44,110,102]
d: [2,169,1,8,2,0,21,168,104]
s: 169
i: [1,114,0,10,3,0,0,255] rc1_tops_1
s: 114
i: [1,115,0,10,3,0,0,255] rc2_tops_1
s: 115
d: [2,170,1,12,4,0,44,112,102]
d: [2,171,1,10,3,0,21,170,104]
s: 171
i: [1,111,0,11,3,0,0,2] t15_tops_1
s: 111
d: [2,172,1,12,4,0,44,113,102]
d: [2,173,1,10,3,0,21,172,104]
s: 173
i: [1,119,0,14,3,0,0,2] dt_k2_pre_topc
s: 119
i: [1,120,0,14,3,0,0,2] dt_k1_tops_1
s: 120
d: [2,174,1,12,4,0,44,114,102]
d: [2,175,1,10,3,0,21,174,104]
s: 175
i: [1,121,0,16,3,0,0,2] projectivity_k1_tops_1
s: 121
d: [2,176,1,12,4,0,44,115,102]
d: [2,177,1,10,3,0,21,176,104]
s: 177
i: [1,122,0,16,3,0,0,2] projectivity_k2_pre_topc
s: 122
i: [1,106,0,17,4,0,0,1] d7_tops_1
s: 106
d: [2,179,1,16,4,0,44,119,102]
d: [2,181,1,14,3,0,21,179,104]
s: 181
i: [1,107,0,17,4,0,0,1] d7_tops_1
s: 107
d: [2,183,1,16,4,0,44,120,102]
d: [2,185,1,14,3,0,21,183,104]
s: 185
i: [1,108,0,17,4,0,0,1] d8_tops_1
s: 108
i: [1,109,0,17,4,0,0,1] d8_tops_1
s: 109
d: [2,214,1,23,6,0,44,107,111]
d: [2,219,1,21,5,0,18,214]
d: [2,220,1,14,4,0,21,219,105]
s: 220
d: [2,200,1,24,6,0,44,122,106]
d: [2,203,1,22,5,0,18,200]
d: [2,207,1,15,4,0,21,203,120]
s: 207
d: [2,253,2,10,4,0,40,220,103]
d: [2,256,2,8,3,0,21,253,104]
s: 256
r: 220
d: [2,229,1,31,7,0,44,107,108]
d: [2,234,1,29,6,0,18,229]
d: [2,235,1,22,5,0,19,234]
d: [2,240,1,15,4,0,21,235,119]
s: 240
d: [2,230,1,24,6,0,44,121,108]
d: [2,233,1,22,5,0,18,230]
d: [2,241,1,15,4,0,21,233,119]
s: 241
d: [2,281,2,14,4,0,40,240,110]
d: [2,292,2,12,3,0,18,281]
s: 292
d: [2,244,1,31,7,0,44,109,106]
d: [2,247,1,29,6,0,18,244]
d: [2,248,1,22,5,0,19,247]
d: [2,252,1,15,4,0,21,248,120]
s: 252
d: [2,278,2,14,4,0,40,240,103]
d: [2,295,2,12,3,0,21,278,104]
s: 295
d: [2,226,1,21,5,0,44,108,105]
d: [2,238,1,19,4,0,18,226]
s: 238
d: [2,242,1,21,5,0,44,109,105]
d: [2,250,1,19,4,0,18,242]
s: 250
d: [2,337,3,12,4,0,44,295,105]
d: [2,338,3,10,3,0,18,337]
s: 338
d: [2,199,1,30,6,0,44,106,121]
d: [2,204,1,28,5,0,18,199]
d: [2,206,1,21,4,0,21,204,120]
s: 206
d: [2,321,2,14,4,0,40,252,110]
d: [2,332,2,12,3,0,18,321]
s: 332
d: [2,215,1,30,6,0,44,107,121]
d: [2,218,1,28,5,0,18,215]
d: [2,221,1,21,4,0,21,218,120]
s: 221
d: [2,227,1,30,6,0,44,108,122]
d: [2,237,1,28,5,0,18,227]
d: [2,239,1,21,4,0,21,237,119]
s: 239
d: [2,318,2,14,4,0,40,252,103]
d: [2,335,2,12,3,0,21,318,104]
s: 335
d: [2,243,1,30,6,0,44,109,122]
d: [2,249,1,28,5,0,18,243]
d: [2,251,1,21,4,0,21,249,119]
s: 251
d: [2,409,3,14,5,0,44,335,111]
d: [2,410,3,12,4,0,18,409]
d: [2,411,3,8,3,0,21,410,256]
s: 411
d: [2,427,4,0,0,3,52,142,137,128,411]
e: 427
d: [2,428,4,4,2,1,40,411,128]
d: [2,432,4,2,1,2,21,428,137]
d: [2,433,4,0,0,3,21,432,142]
e: 433
r: 128
a: [3,127,0,4,1,101]
s: 127
d: [2,435,4,6,2,1,40,127,338]
d: [2,436,4,4,1,2,21,435,137]
f: 436,433,99,100,101
% Refutation found. Thanks to Tanya!
% SZS status Theorem for isomichi__t11_isomichi
% SZS output start Proof for isomichi__t11_isomichi
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => (v5_tops_1(k2_struct_0(X0),X0) & v6_tops_1(k2_struct_0(X0),X0))) [input t11_isomichi]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => (v5_tops_1(k2_struct_0(X0),X0) & v6_tops_1(k2_struct_0(X0),X0))) [negated conjecture 1]
5. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v5_tops_1(X1,X0) <=> k2_pre_topc(X0,k1_tops_1(X0,X1)) = X1))) [input d7_tops_1]
6. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v6_tops_1(X1,X0) <=> k1_tops_1(X0,k2_pre_topc(X0,X1)) = X1))) [input d8_tops_1]
7. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) & l1_pre_topc(X0)) => m1_subset_1(k1_tops_1(X0,X1),k1_zfmisc_1(u1_struct_0(X0)))) [input dt_k1_tops_1]
9. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) & l1_pre_topc(X0)) => m1_subset_1(k2_pre_topc(X0,X1),k1_zfmisc_1(u1_struct_0(X0)))) [input dt_k2_pre_topc]
10. ! [X0] : (l1_struct_0(X0) => m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0)))) [input dt_k2_struct_0]
11. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
32. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => k2_struct_0(X0) = k1_tops_1(X0,k2_struct_0(X0))) [input t15_tops_1]
33. ! [X0] : (l1_pre_topc(X0) => k2_struct_0(X0) = k2_pre_topc(X0,k2_struct_0(X0))) [input t2_tops_1]
51. ? [X0] : ((~v5_tops_1(k2_struct_0(X0),X0) | ~v6_tops_1(k2_struct_0(X0),X0)) & (l1_pre_topc(X0) & v2_pre_topc(X0))) [ennf transformation 2]
52. ? [X0] : ((~v5_tops_1(k2_struct_0(X0),X0) | ~v6_tops_1(k2_struct_0(X0),X0)) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 51]
54. ! [X0] : (m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0)) [ennf transformation 10]
55. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 11]
56. ! [X0] : (k2_struct_0(X0) = k2_pre_topc(X0,k2_struct_0(X0)) | ~l1_pre_topc(X0)) [ennf transformation 33]
57. ! [X0] : (! [X1] : ((v5_tops_1(X1,X0) <=> k2_pre_topc(X0,k1_tops_1(X0,X1)) = X1) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 5]
58. ! [X0] : (! [X1] : ((v6_tops_1(X1,X0) <=> k1_tops_1(X0,k2_pre_topc(X0,X1)) = X1) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 6]
60. ! [X0] : (k2_struct_0(X0) = k1_tops_1(X0,k2_struct_0(X0)) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 32]
61. ! [X0] : (k2_struct_0(X0) = k1_tops_1(X0,k2_struct_0(X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 60]
71. ! [X0,X1] : (m1_subset_1(k2_pre_topc(X0,X1),k1_zfmisc_1(u1_struct_0(X0))) | (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0))) [ennf transformation 9]
72. ! [X0,X1] : (m1_subset_1(k2_pre_topc(X0,X1),k1_zfmisc_1(u1_struct_0(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0)) [flattening 71]
73. ! [X0,X1] : (m1_subset_1(k1_tops_1(X0,X1),k1_zfmisc_1(u1_struct_0(X0))) | (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0))) [ennf transformation 7]
74. ! [X0,X1] : (m1_subset_1(k1_tops_1(X0,X1),k1_zfmisc_1(u1_struct_0(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0)) [flattening 73]
79. ? [X0] : ((~v5_tops_1(k2_struct_0(X0),X0) | ~v6_tops_1(k2_struct_0(X0),X0)) & l1_pre_topc(X0) & v2_pre_topc(X0)) => ((~v5_tops_1(k2_struct_0(sK0),sK0) | ~v6_tops_1(k2_struct_0(sK0),sK0)) & l1_pre_topc(sK0) & v2_pre_topc(sK0)) [choice axiom]
80. (~v5_tops_1(k2_struct_0(sK0),sK0) | ~v6_tops_1(k2_struct_0(sK0),sK0)) & l1_pre_topc(sK0) & v2_pre_topc(sK0) [skolemisation 52,79]
81. ! [X0] : (! [X1] : (((v5_tops_1(X1,X0) | k2_pre_topc(X0,k1_tops_1(X0,X1)) != X1) & (k2_pre_topc(X0,k1_tops_1(X0,X1)) = X1 | ~v5_tops_1(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [nnf transformation 57]
82. ! [X0] : (! [X1] : (((v6_tops_1(X1,X0) | k1_tops_1(X0,k2_pre_topc(X0,X1)) != X1) & (k1_tops_1(X0,k2_pre_topc(X0,X1)) = X1 | ~v6_tops_1(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [nnf transformation 58]
99. v2_pre_topc(sK0) [cnf transformation 80] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
100. l1_pre_topc(sK0) [cnf transformation 80] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
101. ~v5_tops_1(k2_struct_0(sK0),sK0) | ~v6_tops_1(k2_struct_0(sK0),sK0) [cnf transformation 80] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
103. m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) [cnf transformation 54] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
104. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 55] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
105. k2_struct_0(X0) = k2_pre_topc(X0,k2_struct_0(X0)) | ~l1_pre_topc(X0) [cnf transformation 56] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
106. k2_pre_topc(X0,k1_tops_1(X0,X1)) = X1 | ~v5_tops_1(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 81] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
107. k2_pre_topc(X0,k1_tops_1(X0,X1)) != X1 | v5_tops_1(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 81] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
108. k1_tops_1(X0,k2_pre_topc(X0,X1)) = X1 | ~v6_tops_1(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 82] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
109. k1_tops_1(X0,k2_pre_topc(X0,X1)) != X1 | v6_tops_1(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 82] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
111. k2_struct_0(X0) = k1_tops_1(X0,k2_struct_0(X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 61] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
119. m1_subset_1(k2_pre_topc(X0,X1),k1_zfmisc_1(u1_struct_0(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 72] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
120. m1_subset_1(k1_tops_1(X0,X1),k1_zfmisc_1(u1_struct_0(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 74] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
126. 1 <=> v6_tops_1(k2_struct_0(sK0),sK0) [avatar definition]
127. v6_tops_1(k2_struct_0(sK0),sK0) <- (1) [avatar component clause 126] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
128. ~v6_tops_1(k2_struct_0(sK0),sK0) <- (~1) [avatar component clause 126] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
130. 2 <=> v5_tops_1(k2_struct_0(sK0),sK0) [avatar definition]
133. ~1 | ~2 [avatar split clause 101,130,126]
135. 3 <=> l1_pre_topc(sK0) [avatar definition]
137. l1_pre_topc(sK0) <- (3) [avatar component clause 135] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
138. 3 [avatar split clause 100,135]
140. 4 <=> v2_pre_topc(sK0) [avatar definition]
142. v2_pre_topc(sK0) <- (4) [avatar component clause 140] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
143. 4 [avatar split clause 99,140]
214. k2_struct_0(X0) != k2_pre_topc(X0,k2_struct_0(X0)) | v5_tops_1(k2_struct_0(X0),X0) | ~m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [superposition 107,111] {a:1,w:23,wCS:9200,sine:1,thAx:0,allAx:2,thDist:-2}
219. k2_struct_0(X0) != k2_pre_topc(X0,k2_struct_0(X0)) | v5_tops_1(k2_struct_0(X0),X0) | ~m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [duplicate literal removal 214] {a:1,w:21,wCS:8400,sine:1,thAx:0,allAx:2,thDist:-2}
220. ~m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | v5_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [subsumption resolution 219,105] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
229. k2_pre_topc(X2,X3) != k2_pre_topc(X2,X3) | v5_tops_1(k2_pre_topc(X2,X3),X2) | ~m1_subset_1(k2_pre_topc(X2,X3),k1_zfmisc_1(u1_struct_0(X2))) | ~l1_pre_topc(X2) | ~v6_tops_1(X3,X2) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))) | ~l1_pre_topc(X2) [superposition 107,108] {a:1,w:31,wCS:12400,sine:1,thAx:0,allAx:2,thDist:-2}
234. k2_pre_topc(X2,X3) != k2_pre_topc(X2,X3) | v5_tops_1(k2_pre_topc(X2,X3),X2) | ~m1_subset_1(k2_pre_topc(X2,X3),k1_zfmisc_1(u1_struct_0(X2))) | ~l1_pre_topc(X2) | ~v6_tops_1(X3,X2) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))) [duplicate literal removal 229] {a:1,w:29,wCS:11600,sine:1,thAx:0,allAx:2,thDist:-2}
235. v5_tops_1(k2_pre_topc(X2,X3),X2) | ~m1_subset_1(k2_pre_topc(X2,X3),k1_zfmisc_1(u1_struct_0(X2))) | ~l1_pre_topc(X2) | ~v6_tops_1(X3,X2) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))) [trivial inequality removal 234] {a:1,w:22,wCS:8800,sine:1,thAx:0,allAx:2,thDist:-2}
240. ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))) | ~l1_pre_topc(X2) | ~v6_tops_1(X3,X2) | v5_tops_1(k2_pre_topc(X2,X3),X2) [subsumption resolution 235,119] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
244. k1_tops_1(X3,X4) != k1_tops_1(X3,X4) | v6_tops_1(k1_tops_1(X3,X4),X3) | ~m1_subset_1(k1_tops_1(X3,X4),k1_zfmisc_1(u1_struct_0(X3))) | ~l1_pre_topc(X3) | ~v5_tops_1(X4,X3) | ~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))) | ~l1_pre_topc(X3) [superposition 109,106] {a:1,w:31,wCS:12400,sine:1,thAx:0,allAx:2,thDist:-2}
247. k1_tops_1(X3,X4) != k1_tops_1(X3,X4) | v6_tops_1(k1_tops_1(X3,X4),X3) | ~m1_subset_1(k1_tops_1(X3,X4),k1_zfmisc_1(u1_struct_0(X3))) | ~l1_pre_topc(X3) | ~v5_tops_1(X4,X3) | ~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))) [duplicate literal removal 244] {a:1,w:29,wCS:11600,sine:1,thAx:0,allAx:2,thDist:-2}
248. v6_tops_1(k1_tops_1(X3,X4),X3) | ~m1_subset_1(k1_tops_1(X3,X4),k1_zfmisc_1(u1_struct_0(X3))) | ~l1_pre_topc(X3) | ~v5_tops_1(X4,X3) | ~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))) [trivial inequality removal 247] {a:1,w:22,wCS:8800,sine:1,thAx:0,allAx:2,thDist:-2}
252. ~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))) | ~l1_pre_topc(X3) | ~v5_tops_1(X4,X3) | v6_tops_1(k1_tops_1(X3,X4),X3) [subsumption resolution 248,120] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
253. v5_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | ~l1_struct_0(X0) [resolution 220,103] {a:2,w:10,wCS:4000,sine:1,thAx:0,allAx:3,thDist:-3}
256. v5_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [subsumption resolution 253,104] {a:2,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
278. ~l1_pre_topc(X2) | ~v6_tops_1(k2_struct_0(X2),X2) | v5_tops_1(k2_pre_topc(X2,k2_struct_0(X2)),X2) | ~l1_struct_0(X2) [resolution 240,103] {a:2,w:14,wCS:5600,sine:1,thAx:0,allAx:3,thDist:-3}
295. v5_tops_1(k2_pre_topc(X2,k2_struct_0(X2)),X2) | ~v6_tops_1(k2_struct_0(X2),X2) | ~l1_pre_topc(X2) [subsumption resolution 278,104] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
318. ~l1_pre_topc(X2) | ~v5_tops_1(k2_struct_0(X2),X2) | v6_tops_1(k1_tops_1(X2,k2_struct_0(X2)),X2) | ~l1_struct_0(X2) [resolution 252,103] {a:2,w:14,wCS:5600,sine:1,thAx:0,allAx:3,thDist:-3}
335. v6_tops_1(k1_tops_1(X2,k2_struct_0(X2)),X2) | ~v5_tops_1(k2_struct_0(X2),X2) | ~l1_pre_topc(X2) [subsumption resolution 318,104] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
337. v5_tops_1(k2_struct_0(X0),X0) | ~v6_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~l1_pre_topc(X0) [superposition 295,105] {a:3,w:12,wCS:4800,sine:1,thAx:0,allAx:4,thDist:-4}
338. ~v6_tops_1(k2_struct_0(X0),X0) | v5_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) [duplicate literal removal 337] {a:3,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
409. v6_tops_1(k2_struct_0(X0),X0) | ~v5_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [superposition 335,111] {a:3,w:14,wCS:5600,sine:1,thAx:0,allAx:4,thDist:-4}
410. v6_tops_1(k2_struct_0(X0),X0) | ~v5_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [duplicate literal removal 409] {a:3,w:12,wCS:4800,sine:1,thAx:0,allAx:4,thDist:-4}
411. v6_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [subsumption resolution 410,256] {a:3,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
428. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~1) [resolution 411,128] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
432. ~v2_pre_topc(sK0) <- (~1, 3) [subsumption resolution 428,137] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
433. $false <- (~1, 3, 4) [subsumption resolution 432,142] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
434. 1 | ~3 | ~4 [avatar contradiction clause 433]
435. v5_tops_1(k2_struct_0(sK0),sK0) | ~l1_pre_topc(sK0) <- (1) [resolution 127,338] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
436. v5_tops_1(k2_struct_0(sK0),sK0) <- (1, 3) [subsumption resolution 435,137] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
437. 2 | ~1 | ~3 [avatar split clause 436,135,126,130]
438. $false [avatar sat refutation 133,138,143,434,437] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:51,thDist:-51}
% SZS output end Proof for isomichi__t11_isomichi
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 33
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 5
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 306
% Active clauses: 54
% Passive clauses: 112
% Final active clauses: 52
% Final passive clauses: 57
% Main loop iterations started: 55
% 
% >>> Simplifying Inferences
% Duplicate literals: 111
% Trivial inequalities: 3
% Fw subsumption resolutions: 39
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Forward subsumptions: 26
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 47
% Unit resulting resolution: 1
% Forward superposition: 56
% Backward superposition: 24
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 4
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 4989
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

