i: [1,183,0,2,1,1,0,0] t5_xxreal_1
a: [3,302,0,2,1,183]
s: 302
i: [1,182,0,2,1,1,0,0] t5_xxreal_1
a: [3,307,0,2,1,182]
s: 307
i: [1,187,0,3,1,1,0,0] t5_xxreal_1
a: [3,282,0,3,1,187]
s: 282
i: [1,186,0,3,1,1,0,0] t5_xxreal_1
a: [3,287,0,3,1,186]
s: 287
i: [1,185,0,5,1,1,0,0] t5_xxreal_1
a: [3,292,0,5,1,185]
s: 292
i: [1,184,0,5,1,1,0,0] t5_xxreal_1
a: [3,297,0,5,1,184]
s: 297
i: [1,188,0,2,1,0,0,3] fc5_numbers
a: [3,312,0,2,1,188]
s: 312
i: [1,189,0,2,1,0,0,2] fc2_membered
a: [3,317,0,2,1,189]
s: 317
i: [1,190,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,322,0,2,1,190]
s: 322
i: [1,191,0,2,1,0,0,255] fc1_xboole_0
a: [3,327,0,2,1,191]
s: 327
i: [1,252,0,2,1,0,0,2] rc2_xboole_0
a: [3,344,0,2,1,252]
s: 344
i: [1,253,0,2,1,0,0,255] rc1_xreal_0
a: [3,349,0,2,1,253]
s: 349
i: [1,254,0,2,1,0,0,2] rc1_xboole_0
a: [3,354,0,2,1,254]
s: 354
i: [1,255,0,2,1,0,0,1] rc1_xxreal_0
a: [3,359,0,2,1,255]
s: 359
i: [1,257,0,2,1,0,0,255] rc1_membered
a: [3,364,0,2,1,257]
s: 364
i: [1,256,0,2,1,0,0,255] rc1_membered
a: [3,369,0,2,1,256]
s: 369
i: [1,259,0,2,1,0,0,255] rc2_membered
a: [3,374,0,2,1,259]
s: 374
i: [1,258,0,2,1,0,0,255] rc2_membered
a: [3,379,0,2,1,258]
s: 379
i: [1,261,0,2,1,0,0,255] rc3_membered
a: [3,384,0,2,1,261]
s: 384
i: [1,260,0,2,1,0,0,255] rc3_membered
a: [3,389,0,2,1,260]
s: 389
i: [1,264,0,2,1,0,0,255] rc3_xreal_0
a: [3,394,0,2,1,264]
s: 394
i: [1,263,0,2,1,0,0,255] rc3_xreal_0
a: [3,399,0,2,1,263]
s: 399
i: [1,262,0,2,1,0,0,255] rc3_xreal_0
a: [3,404,0,2,1,262]
s: 404
i: [1,267,0,2,1,0,0,255] rc2_xreal_0
a: [3,409,0,2,1,267]
s: 409
i: [1,266,0,2,1,0,0,255] rc2_xreal_0
a: [3,414,0,2,1,266]
s: 414
i: [1,265,0,2,1,0,0,255] rc2_xreal_0
a: [3,419,0,2,1,265]
s: 419
i: [1,270,0,2,1,0,0,255] rc4_xreal_0
a: [3,424,0,2,1,270]
s: 424
i: [1,269,0,2,1,0,0,255] rc4_xreal_0
a: [3,429,0,2,1,269]
s: 429
i: [1,268,0,2,1,0,0,255] rc4_xreal_0
a: [3,434,0,2,1,268]
s: 434
i: [1,272,0,2,1,0,0,2] rc4_xxreal_0
a: [3,439,0,2,1,272]
s: 439
i: [1,271,0,2,1,0,0,2] rc4_xxreal_0
a: [3,444,0,2,1,271]
s: 444
i: [1,274,0,2,1,0,0,255] rc3_xxreal_0
a: [3,449,0,2,1,274]
s: 449
i: [1,273,0,2,1,0,0,255] rc3_xxreal_0
a: [3,454,0,2,1,273]
s: 454
i: [1,276,0,2,1,0,0,255] rc2_xxreal_0
a: [3,459,0,2,1,276]
s: 459
i: [1,275,0,2,1,0,0,255] rc2_xxreal_0
a: [3,464,0,2,1,275]
s: 464
i: [1,192,0,3,1,0,0,255] d2_xboole_0
a: [3,332,0,3,1,192]
s: 332
r: 322
i: [1,193,0,4,2,0,0,255] cc2_membered
s: 193
i: [1,194,0,4,2,0,0,255] cc3_membered
s: 194
i: [1,195,0,4,2,0,0,255] cc2_xreal_0
s: 195
i: [1,196,0,4,2,0,0,255] cc2_xxreal_0
s: 196
i: [1,197,0,4,2,0,0,255] cc4_membered
s: 197
i: [1,200,0,4,2,0,0,255] cc1_membered
s: 200
i: [1,202,0,4,2,0,0,255] cc4_xreal_0
s: 202
d: [2,467,1,2,1,1,40,200,364]
a: [3,473,1,2,1,467]
s: 473
i: [1,203,0,4,2,0,0,255] cc18_membered
s: 203
i: [1,236,0,4,1,0,0,2] existence_m1_subset_1
s: 236
d: [2,468,1,2,1,1,40,200,374]
a: [3,478,1,2,1,468]
s: 478
i: [1,204,0,5,2,0,0,255] t6_boole
s: 204
d: [2,469,1,2,1,1,40,200,384]
a: [3,483,1,2,1,469]
s: 483
i: [1,214,0,5,2,0,0,3] cc1_xxreal_0
s: 214
i: [1,233,0,5,2,0,0,1] t7_boole
s: 233
d: [2,544,1,4,1,1,40,233,297]
a: [3,548,1,4,1,544]
s: 548
i: [1,238,0,7,3,0,0,2] reflexivity_r1_xxreal_0
a: [3,339,0,5,2,238]
s: 339
d: [2,485,1,2,1,1,40,202,349]
a: [3,492,1,2,1,485]
s: 492
i: [1,220,0,6,3,0,0,255] cc7_xxreal_0
s: 220
i: [1,221,0,6,3,0,0,255] cc7_xxreal_0
s: 221
d: [2,500,1,2,1,1,40,203,327]
a: [3,507,1,2,1,500]
s: 507
i: [1,224,0,6,3,0,0,255] cc5_xxreal_0
s: 224
d: [2,501,1,2,1,1,40,203,354]
a: [3,512,1,2,1,501]
s: 512
i: [1,226,0,6,3,0,0,255] cc5_xxreal_0
s: 226
i: [1,227,0,6,3,0,0,255] cc3_xxreal_0
s: 227
d: [2,502,1,2,1,1,40,203,434]
a: [3,517,1,2,1,502]
s: 517
i: [1,229,0,6,3,0,0,255] cc3_xxreal_0
s: 229
d: [2,503,1,2,1,1,40,203,444]
a: [3,522,1,2,1,503]
s: 522
i: [1,231,0,6,2,0,0,1] t1_subset
s: 231
i: [1,237,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 237
d: [2,617,1,5,1,1,40,231,297]
a: [3,621,1,5,1,617]
s: 621
i: [1,198,0,7,3,0,0,255] cc14_membered
s: 198
d: [2,630,2,6,2,1,40,198,621]
a: [3,634,2,4,1,630]
s: 634
i: [1,199,0,7,3,0,0,2] cc13_membered
s: 199
i: [1,201,0,7,3,0,0,255] cc17_membered
s: 201
d: [2,641,2,6,2,1,40,199,621]
a: [3,645,2,4,1,641]
s: 645
i: [1,232,0,7,3,0,0,1] t8_boole
s: 232
d: [2,652,2,6,2,1,40,201,621]
a: [3,656,2,4,1,652]
s: 656
i: [1,234,0,7,2,0,0,1] t2_tarski
s: 234
i: [1,216,0,8,4,0,0,255] cc6_xxreal_0
s: 216
d: [2,665,1,6,3,1,40,216,307]
a: [3,678,1,2,1,665]
s: 678
r: 678
i: [1,218,0,8,4,0,0,255] cc4_xxreal_0
s: 218
a: [3,677,1,2,1,665]
s: 677
i: [1,222,0,8,4,0,0,255] cc8_xxreal_0
s: 222
i: [1,230,0,8,3,0,0,1] t2_subset
s: 230
a: [3,682,1,2,1,665]
s: 682
i: [1,239,0,8,3,0,0,1] fc1_xxreal_1
s: 239
d: [2,739,3,0,0,3,52,307,302,645,239]
e: 739
d: [2,740,3,4,2,1,40,239,645]
d: [2,742,3,2,1,2,21,740,302]
d: [2,743,3,0,0,3,21,742,307]
e: 743
r: 645
a: [3,685,1,2,1,665]
s: 685
i: [1,240,0,8,3,0,0,1] fc4_xxreal_1
s: 240
i: [1,235,0,10,3,0,0,1] t2_tarski
s: 235
d: [2,666,1,6,3,1,40,216,302]
a: [3,691,1,2,1,666]
s: 691
r: 691
i: [1,241,0,10,4,0,0,2] connectedness_r1_xxreal_0
s: 241
a: [3,690,1,2,1,666]
s: 690
i: [1,205,0,11,3,0,0,255] d4_xxreal_1
s: 205
i: [1,206,0,11,3,0,0,255] d1_xxreal_1
s: 206
a: [3,695,1,2,1,666]
s: 695
i: [1,207,0,13,5,0,0,2] t1_xxreal_0
s: 207
a: [3,698,1,2,1,666]
s: 698
i: [1,208,0,14,5,0,0,1] t4_xxreal_1
s: 208
i: [1,209,0,14,5,0,0,1] t4_xxreal_1
s: 209
a: [3,649,2,2,1,641]
s: 649
i: [1,211,0,14,5,0,0,1] t1_xxreal_1
s: 211
d: [2,764,3,6,3,1,40,649,216]
a: [3,768,3,2,1,764]
s: 768
r: 768
i: [1,212,0,14,5,0,0,1] t1_xxreal_1
s: 212
i: [1,242,0,15,4,0,0,255] fraenkel_a_2_3_xxreal_1
s: 242
a: [3,767,3,2,1,764]
s: 767
i: [1,243,0,15,4,0,0,255] fraenkel_a_2_3_xxreal_1
s: 243
a: [3,772,3,2,1,764]
s: 772
i: [1,244,0,15,4,0,0,255] fraenkel_a_2_3_xxreal_1
s: 244
i: [1,245,0,15,4,0,0,255] fraenkel_a_2_3_xxreal_1
s: 245
a: [3,775,3,2,1,764]
s: 775
i: [1,247,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 247
d: [2,778,1,9,4,1,40,211,297]
d: [2,780,1,7,3,2,21,778,302]
d: [2,781,1,5,2,3,21,780,307]
d: [2,782,1,3,1,4,21,781,649]
a: [3,786,1,3,1,782]
s: 786
i: [1,248,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 248
i: [1,249,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 249
d: [2,794,1,9,4,1,40,212,297]
d: [2,796,1,7,3,2,21,794,302]
d: [2,797,1,5,2,3,21,796,307]
d: [2,798,1,3,1,4,21,797,649]
a: [3,802,1,3,1,798]
s: 802
i: [1,250,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 250
d: [2,856,2,10,4,1,40,786,207]
d: [2,857,2,7,3,2,21,856,287]
d: [2,858,2,5,2,3,21,857,307]
d: [2,859,2,3,1,4,21,858,649]
a: [3,863,2,3,1,859]
s: 863
i: [1,210,0,17,6,0,0,1] t4_xxreal_1
s: 210
d: [2,871,2,10,4,1,40,802,207]
d: [2,872,2,7,3,2,21,871,282]
d: [2,873,2,5,2,3,21,872,649]
d: [2,874,2,3,1,4,21,873,302]
d: [2,885,1,12,5,1,40,210,292]
d: [2,893,1,9,4,2,21,885,863]
d: [2,894,1,7,3,3,21,893,302]
d: [2,895,1,5,2,4,21,894,307]
d: [2,896,1,3,1,5,21,895,649]
f: 896,874,859,798,782,743,641,617,182,183,184,185,186,187
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_1__t5_xxreal_1
% SZS output start Proof for xxreal_1__t5_xxreal_1
1. ! [X0] : ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ~(X0 != X2 & X0 != X1 & ~r2_hidden(X0,k4_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2))))) [input t5_xxreal_1]
2. ~! [X0] : ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ~(X0 != X2 & X0 != X1 & ~r2_hidden(X0,k4_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2))))) [negated conjecture 1]
5. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xxreal_0(X1))) [input cc13_membered]
40. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v2_membered(k1_xxreal_1(X0,X1))) [input fc1_xxreal_1]
60. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
61. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input t1_xxreal_0]
62. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k1_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0)))))) [input t1_xxreal_1]
65. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k4_xxreal_1(X1,X2)) <=> (~r1_xxreal_0(X2,X0) & ~r1_xxreal_0(X0,X1)))))) [input t4_xxreal_1]
69. ~! [X0,X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ~(X0 != X2 & X0 != X1 & ~r2_hidden(X0,k4_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2))))) [flattening 2]
71. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 60]
85. ? [X0,X1] : (? [X2] : ((X0 != X2 & X0 != X1 & ~r2_hidden(X0,k4_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2))) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) [ennf transformation 69]
86. ? [X0,X1] : (? [X2] : (X0 != X2 & X0 != X1 & ~r2_hidden(X0,k4_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) [flattening 85]
93. ! [X0] : (! [X1] : (v1_xxreal_0(X1) | ~m1_subset_1(X1,X0)) | ~v2_membered(X0)) [ennf transformation 5]
101. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 61]
102. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 101]
103. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X0,k4_xxreal_1(X1,X2)) <=> (~r1_xxreal_0(X2,X0) & ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 65]
104. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X0,k1_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 62]
120. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 71]
127. ! [X0,X1] : (v2_membered(k1_xxreal_1(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 40]
128. ! [X0,X1] : (v2_membered(k1_xxreal_1(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 127]
137. ? [X0,X1] : (? [X2] : (X0 != X2 & X0 != X1 & ~r2_hidden(X0,k4_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (sK0 != X2 & sK0 != sK1 & ~r2_hidden(sK0,k4_xxreal_1(sK1,X2)) & r2_hidden(sK0,k1_xxreal_1(sK1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
138. ? [X2] : (sK0 != X2 & sK0 != sK1 & ~r2_hidden(sK0,k4_xxreal_1(sK1,X2)) & r2_hidden(sK0,k1_xxreal_1(sK1,X2)) & v1_xxreal_0(X2)) => (sK0 != sK2 & sK0 != sK1 & ~r2_hidden(sK0,k4_xxreal_1(sK1,sK2)) & r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) & v1_xxreal_0(sK2)) [choice axiom]
139. (sK0 != sK2 & sK0 != sK1 & ~r2_hidden(sK0,k4_xxreal_1(sK1,sK2)) & r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1) [skolemisation 86,138,137]
140. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k4_xxreal_1(X1,X2)) | (r1_xxreal_0(X2,X0) | r1_xxreal_0(X0,X1))) & ((~r1_xxreal_0(X2,X0) & ~r1_xxreal_0(X0,X1)) | ~r2_hidden(X0,k4_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 103]
141. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k4_xxreal_1(X1,X2)) | r1_xxreal_0(X2,X0) | r1_xxreal_0(X0,X1)) & ((~r1_xxreal_0(X2,X0) & ~r1_xxreal_0(X0,X1)) | ~r2_hidden(X0,k4_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 140]
142. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k1_xxreal_1(X1,X2)) | (~r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X0))) & ((r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0)) | ~r2_hidden(X0,k1_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 104]
143. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k1_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X0)) & ((r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0)) | ~r2_hidden(X0,k1_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 142]
182. v1_xxreal_0(sK1) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
183. v1_xxreal_0(sK2) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
184. r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) [cnf transformation 139] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
185. ~r2_hidden(sK0,k4_xxreal_1(sK1,sK2)) [cnf transformation 139] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
186. sK0 != sK1 [cnf transformation 139] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
187. sK0 != sK2 [cnf transformation 139] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
199. ~m1_subset_1(X1,X0) | v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 93] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
207. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 102] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
210. r2_hidden(X0,k4_xxreal_1(X1,X2)) | r1_xxreal_0(X2,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 141] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
211. ~r2_hidden(X0,k1_xxreal_1(X1,X2)) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 143] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
212. ~r2_hidden(X0,k1_xxreal_1(X1,X2)) | r1_xxreal_0(X0,X2) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 143] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
231. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 120] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
239. v2_membered(k1_xxreal_1(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 128] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
280. 1 <=> sK0 = sK2 [avatar definition]
282. sK0 != sK2 <- (~1) [avatar component clause 280] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
283. ~1 [avatar split clause 187,280]
285. 2 <=> sK0 = sK1 [avatar definition]
287. sK0 != sK1 <- (~2) [avatar component clause 285] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
288. ~2 [avatar split clause 186,285]
290. 3 <=> r2_hidden(sK0,k4_xxreal_1(sK1,sK2)) [avatar definition]
292. ~r2_hidden(sK0,k4_xxreal_1(sK1,sK2)) <- (~3) [avatar component clause 290] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
293. ~3 [avatar split clause 185,290]
295. 4 <=> r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) [avatar definition]
297. r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) <- (4) [avatar component clause 295] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
298. 4 [avatar split clause 184,295]
300. 5 <=> v1_xxreal_0(sK2) [avatar definition]
302. v1_xxreal_0(sK2) <- (5) [avatar component clause 300] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
303. 5 [avatar split clause 183,300]
305. 6 <=> v1_xxreal_0(sK1) [avatar definition]
307. v1_xxreal_0(sK1) <- (6) [avatar component clause 305] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
308. 6 [avatar split clause 182,305]
617. m1_subset_1(sK0,k1_xxreal_1(sK1,sK2)) <- (4) [resolution 231,297] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
619. 62 <=> m1_subset_1(sK0,k1_xxreal_1(sK1,sK2)) [avatar definition]
621. m1_subset_1(sK0,k1_xxreal_1(sK1,sK2)) <- (62) [avatar component clause 619] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
622. 62 | ~4 [avatar split clause 617,295,619]
641. v1_xxreal_0(sK0) | ~v2_membered(k1_xxreal_1(sK1,sK2)) <- (62) [resolution 199,621] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
643. 66 <=> v2_membered(k1_xxreal_1(sK1,sK2)) [avatar definition]
645. ~v2_membered(k1_xxreal_1(sK1,sK2)) <- (~66) [avatar component clause 643] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
647. 67 <=> v1_xxreal_0(sK0) [avatar definition]
649. v1_xxreal_0(sK0) <- (67) [avatar component clause 647] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
650. ~66 | 67 | ~62 [avatar split clause 641,619,647,643]
740. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~66) [resolution 239,645] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
742. ~v1_xxreal_0(sK1) <- (5, ~66) [subsumption resolution 740,302] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
743. $false <- (5, 6, ~66) [subsumption resolution 742,307] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
744. ~5 | ~6 | 66 [avatar contradiction clause 743]
778. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4) [resolution 211,297] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
780. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4, 5) [subsumption resolution 778,302] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
781. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) <- (4, 5, 6) [subsumption resolution 780,307] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
782. r1_xxreal_0(sK1,sK0) <- (4, 5, 6, 67) [subsumption resolution 781,649] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
784. 87 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
786. r1_xxreal_0(sK1,sK0) <- (87) [avatar component clause 784] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
787. 87 | ~4 | ~5 | ~6 | ~67 [avatar split clause 782,647,305,300,295,784]
794. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4) [resolution 212,297] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
796. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4, 5) [subsumption resolution 794,302] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
797. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK0) <- (4, 5, 6) [subsumption resolution 796,307] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
798. r1_xxreal_0(sK0,sK2) <- (4, 5, 6, 67) [subsumption resolution 797,649] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
800. 89 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
802. r1_xxreal_0(sK0,sK2) <- (89) [avatar component clause 800] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
803. 89 | ~4 | ~5 | ~6 | ~67 [avatar split clause 798,647,305,300,295,800]
856. sK0 = sK1 | ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (87) [resolution 786,207] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
857. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, 87) [subsumption resolution 856,287] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
858. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK0) <- (~2, 6, 87) [subsumption resolution 857,307] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
859. ~r1_xxreal_0(sK0,sK1) <- (~2, 6, 67, 87) [subsumption resolution 858,649] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
861. 96 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
863. ~r1_xxreal_0(sK0,sK1) <- (~96) [avatar component clause 861] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
864. ~96 | 2 | ~6 | ~67 | ~87 [avatar split clause 859,784,647,305,285,861]
871. sK0 = sK2 | ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) <- (89) [resolution 802,207] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
872. ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) <- (~1, 89) [subsumption resolution 871,282] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
873. ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) <- (~1, 67, 89) [subsumption resolution 872,649] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
874. ~r1_xxreal_0(sK2,sK0) <- (~1, 5, 67, 89) [subsumption resolution 873,302] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
876. 97 <=> r1_xxreal_0(sK2,sK0) [avatar definition]
879. ~97 | 1 | ~5 | ~67 | ~89 [avatar split clause 874,800,647,300,280,876]
885. r1_xxreal_0(sK2,sK0) | r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~3) [resolution 210,292] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
893. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~3, ~96) [subsumption resolution 885,863] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
894. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~3, 5, ~96) [subsumption resolution 893,302] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
895. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) <- (~3, 5, 6, ~96) [subsumption resolution 894,307] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
896. r1_xxreal_0(sK2,sK0) <- (~3, 5, 6, 67, ~96) [subsumption resolution 895,649] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
897. 97 | 3 | ~5 | ~6 | ~67 | 96 [avatar split clause 896,861,647,305,300,290,876]
898. $false [avatar sat refutation 283,288,293,298,303,308,622,650,744,787,803,864,879,897] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:450,thDist:-450}
% SZS output end Proof for xxreal_1__t5_xxreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 67
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 5
% 
% >>> Saturation
% Initial clauses: 95
% Generated clauses: 332
% Active clauses: 110
% Passive clauses: 170
% Final active clauses: 105
% Final passive clauses: 55
% Main loop iterations started: 115
% 
% >>> Simplifying Inferences
% Duplicate literals: 33
% Fw subsumption resolutions: 31
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 1
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 92
% Unit resulting resolution: 1
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 13
% Split components: 31
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 89
% SAT solver unit clauses: 37
% SAT solver binary clauses: 26
% 
% Memory used [KB]: 9978
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% forward subsumption resolution: 0.001 s
% superposition: 0.002 s
% other: 0.005 s ( own 0.001 s ) 

