i: [1,81,0,2,1,1,0,0] t89_mcart_1
a: [3,145,0,2,1,81]
s: 145
i: [1,86,0,3,1,1,0,0] t89_mcart_1
a: [3,120,0,3,1,86]
s: 120
i: [1,83,0,3,1,1,0,0] t89_mcart_1
a: [3,135,0,3,1,83]
s: 135
i: [1,82,0,3,1,1,0,0] t89_mcart_1
a: [3,140,0,3,1,82]
s: 140
i: [1,85,0,5,1,1,0,0] t89_mcart_1
a: [3,125,0,5,1,85]
s: 125
i: [1,84,0,5,1,1,0,0] t89_mcart_1
a: [3,130,0,5,1,84]
s: 130
i: [1,87,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,150,0,2,1,87]
s: 150
i: [1,88,0,2,1,0,0,255] fc1_xboole_0
a: [3,155,0,2,1,88]
s: 155
i: [1,111,0,2,1,0,0,2] rc2_xboole_0
a: [3,165,0,2,1,111]
s: 165
i: [1,112,0,2,1,0,0,255] rc1_xtuple_0
a: [3,170,0,2,1,112]
s: 170
i: [1,113,0,2,1,0,0,2] rc1_xboole_0
a: [3,175,0,2,1,113]
s: 175
i: [1,115,0,2,1,0,0,1] rc1_relat_1
a: [3,180,0,2,1,115]
s: 180
i: [1,114,0,2,1,0,0,1] rc1_relat_1
a: [3,185,0,2,1,114]
s: 185
i: [1,116,0,2,1,0,0,255] rc2_relat_1
a: [3,190,0,2,1,116]
s: 190
i: [1,90,0,3,1,0,0,255] fc2_xboole_0
s: 90
i: [1,89,0,3,1,0,0,255] d2_xboole_0
a: [3,160,0,3,1,89]
s: 160
r: 150
i: [1,92,0,4,2,0,0,1] cc1_relat_1
s: 92
i: [1,102,0,4,1,0,0,255] existence_m1_subset_1
s: 102
d: [2,193,1,2,1,1,40,92,175]
a: [3,198,1,2,1,193]
s: 198
i: [1,103,0,4,1,0,0,255] fc3_xboole_0
s: 103
i: [1,104,0,4,1,0,0,255] fc1_xtuple_0
s: 104
d: [2,194,1,2,1,1,40,92,155]
a: [3,203,1,2,1,194]
s: 203
i: [1,93,0,5,2,0,0,255] t6_boole
s: 93
i: [1,101,0,5,2,0,0,1] t7_boole
s: 101
i: [1,105,0,5,1,0,0,255] fc5_relat_1
s: 105
i: [1,99,0,6,2,0,0,255] t1_subset
s: 99
d: [2,208,1,2,1,1,40,101,135]
a: [3,213,1,2,1,208]
s: 213
i: [1,106,0,6,1,0,0,1] rd2_xtuple_0
s: 106
d: [2,216,1,3,1,1,40,99,135]
a: [3,221,1,3,1,216]
s: 221
i: [1,107,0,6,1,0,0,1] rd1_xtuple_0
s: 107
i: [1,109,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 109
d: [2,217,1,3,1,1,40,99,140]
a: [3,226,1,3,1,217]
s: 226
i: [1,100,0,7,3,0,0,1] t8_boole
s: 100
d: [2,228,1,3,1,1,40,109,135]
a: [3,233,1,3,1,228]
s: 233
i: [1,108,0,7,1,0,0,255] commutativity_k2_tarski
s: 108
i: [1,98,0,8,3,0,0,255] t2_subset
s: 98
d: [2,229,1,3,1,1,40,109,140]
a: [3,238,1,3,1,229]
s: 238
i: [1,110,0,8,1,0,0,255] fc7_relat_1
s: 110
d: [2,242,1,6,2,0,40,98,102]
s: 242
i: [1,91,0,9,2,0,0,255] rd3_xtuple_0
s: 91
i: [1,96,0,10,1,0,0,255] d5_tarski
s: 96
d: [2,250,1,9,2,1,44,91,130]
d: [2,258,1,9,2,2,22,250,125]
a: [3,262,1,2,1,258]
s: 262
i: [1,97,0,12,3,0,0,1] t21_mcart_1
s: 97
d: [2,280,1,12,3,1,44,97,130]
d: [2,288,1,12,3,2,22,280,125]
a: [3,291,1,5,2,288]
s: 291
d: [2,295,2,0,0,3,52,145,140,291]
e: 295
d: [2,296,2,2,1,2,40,291,140]
d: [2,298,2,0,0,3,21,296,145]
e: 298
r: 291
a: [3,266,1,7,1,258]
s: 266
d: [2,309,2,2,1,1,44,104,266]
d: [2,326,2,0,0,2,21,309,262]
e: 326
r: 262
a: [3,261,1,2,1,258]
s: 261
d: [2,301,2,5,2,1,44,266,91]
d: [2,311,2,2,1,2,21,301,120]
a: [3,315,2,2,1,311]
s: 315
d: [2,279,1,4,1,0,44,103,96]
s: 279
d: [2,331,2,2,1,1,44,279,266]
a: [3,335,2,2,1,331]
s: 335
d: [2,256,1,5,2,0,44,105,91]
s: 256
d: [2,205,1,7,3,0,44,93,93]
s: 205
d: [2,308,2,3,1,1,44,105,266]
a: [3,324,2,3,1,308]
s: 324
d: [2,287,1,7,3,0,44,104,97]
s: 287
d: [2,337,3,0,0,3,52,145,315,135,287]
e: 337
d: [2,338,2,4,2,1,40,287,135]
d: [2,342,2,2,1,2,21,338,315]
d: [2,343,2,0,0,3,21,342,145]
e: 343
d: [2,282,1,12,3,1,44,97,125]
d: [2,293,1,12,3,2,22,282,130]
d: [2,303,2,5,2,1,44,91,266]
d: [2,318,2,2,1,2,21,303,120]
f: 343,318,298,293,81,82,83,84,85,86
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t89_mcart_1
% SZS output start Proof for mcart_1__t89_mcart_1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : ((k2_xtuple_0(X2) = k2_xtuple_0(X0) & k1_xtuple_0(X2) = k1_xtuple_0(X0) & r2_hidden(X0,X1) & r2_hidden(X2,X1)) => X0 = X2)) [input t89_mcart_1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : ((k2_xtuple_0(X2) = k2_xtuple_0(X0) & k1_xtuple_0(X2) = k1_xtuple_0(X0) & r2_hidden(X0,X1) & r2_hidden(X2,X1)) => X0 = X2)) [negated conjecture 1]
20. ! [X0,X1] : v1_xtuple_0(k4_tarski(X0,X1)) [input fc1_xtuple_0]
32. ! [X0] : (v1_xtuple_0(X0) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0) [input rd3_xtuple_0]
34. ! [X0] : ! [X1] : (v1_relat_1(X1) => (r2_hidden(X0,X1) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0)) [input t21_mcart_1]
39. ~! [X0,X1] : (v1_relat_1(X1) => ! [X2] : ((k2_xtuple_0(X2) = k2_xtuple_0(X0) & k1_xtuple_0(X2) = k1_xtuple_0(X0) & r2_hidden(X0,X1) & r2_hidden(X2,X1)) => X0 = X2)) [flattening 2]
41. ! [X0,X1] : (v1_relat_1(X1) => (r2_hidden(X0,X1) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0)) [flattening 34]
49. ? [X0,X1] : (? [X2] : (X0 != X2 & (k2_xtuple_0(X2) = k2_xtuple_0(X0) & k1_xtuple_0(X2) = k1_xtuple_0(X0) & r2_hidden(X0,X1) & r2_hidden(X2,X1))) & v1_relat_1(X1)) [ennf transformation 39]
50. ? [X0,X1] : (? [X2] : (X0 != X2 & k2_xtuple_0(X2) = k2_xtuple_0(X0) & k1_xtuple_0(X2) = k1_xtuple_0(X0) & r2_hidden(X0,X1) & r2_hidden(X2,X1)) & v1_relat_1(X1)) [flattening 49]
51. ! [X0] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0)) [ennf transformation 32]
58. ! [X0,X1] : ((k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~r2_hidden(X0,X1)) | ~v1_relat_1(X1)) [ennf transformation 41]
59. ! [X0,X1] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~r2_hidden(X0,X1) | ~v1_relat_1(X1)) [flattening 58]
66. ? [X0,X1] : (? [X2] : (X0 != X2 & k2_xtuple_0(X2) = k2_xtuple_0(X0) & k1_xtuple_0(X2) = k1_xtuple_0(X0) & r2_hidden(X0,X1) & r2_hidden(X2,X1)) & v1_relat_1(X1)) => (? [X2] : (sK0 != X2 & k2_xtuple_0(X2) = k2_xtuple_0(sK0) & k1_xtuple_0(X2) = k1_xtuple_0(sK0) & r2_hidden(sK0,sK1) & r2_hidden(X2,sK1)) & v1_relat_1(sK1)) [choice axiom]
67. ? [X2] : (sK0 != X2 & k2_xtuple_0(X2) = k2_xtuple_0(sK0) & k1_xtuple_0(X2) = k1_xtuple_0(sK0) & r2_hidden(sK0,sK1) & r2_hidden(X2,sK1)) => (sK0 != sK2 & k2_xtuple_0(sK0) = k2_xtuple_0(sK2) & k1_xtuple_0(sK0) = k1_xtuple_0(sK2) & r2_hidden(sK0,sK1) & r2_hidden(sK2,sK1)) [choice axiom]
68. (sK0 != sK2 & k2_xtuple_0(sK0) = k2_xtuple_0(sK2) & k1_xtuple_0(sK0) = k1_xtuple_0(sK2) & r2_hidden(sK0,sK1) & r2_hidden(sK2,sK1)) & v1_relat_1(sK1) [skolemisation 50,67,66]
81. v1_relat_1(sK1) [cnf transformation 68] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
82. r2_hidden(sK2,sK1) [cnf transformation 68] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
83. r2_hidden(sK0,sK1) [cnf transformation 68] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
84. k1_xtuple_0(sK0) = k1_xtuple_0(sK2) [cnf transformation 68] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
85. k2_xtuple_0(sK0) = k2_xtuple_0(sK2) [cnf transformation 68] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
86. sK0 != sK2 [cnf transformation 68] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
91. k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0) [cnf transformation 51] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
97. k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~r2_hidden(X0,X1) | ~v1_relat_1(X1) [cnf transformation 59] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. v1_xtuple_0(k4_tarski(X0,X1)) [cnf transformation 20] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
118. 1 <=> sK0 = sK2 [avatar definition]
120. sK0 != sK2 <- (~1) [avatar component clause 118] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
121. ~1 [avatar split clause 86,118]
123. 2 <=> k2_xtuple_0(sK0) = k2_xtuple_0(sK2) [avatar definition]
125. k2_xtuple_0(sK0) = k2_xtuple_0(sK2) <- (2) [avatar component clause 123] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. 2 [avatar split clause 85,123]
128. 3 <=> k1_xtuple_0(sK0) = k1_xtuple_0(sK2) [avatar definition]
130. k1_xtuple_0(sK0) = k1_xtuple_0(sK2) <- (3) [avatar component clause 128] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
131. 3 [avatar split clause 84,128]
133. 4 <=> r2_hidden(sK0,sK1) [avatar definition]
135. r2_hidden(sK0,sK1) <- (4) [avatar component clause 133] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
136. 4 [avatar split clause 83,133]
138. 5 <=> r2_hidden(sK2,sK1) [avatar definition]
140. r2_hidden(sK2,sK1) <- (5) [avatar component clause 138] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
141. 5 [avatar split clause 82,138]
143. 6 <=> v1_relat_1(sK1) [avatar definition]
145. v1_relat_1(sK1) <- (6) [avatar component clause 143] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
146. 6 [avatar split clause 81,143]
264. 24 <=> sK2 = k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) [avatar definition]
266. sK2 = k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (24) [avatar component clause 264] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
282. sK2 = k4_tarski(k1_xtuple_0(sK2),k2_xtuple_0(sK0)) | ~r2_hidden(sK2,X0) | ~v1_relat_1(X0) <- (2) [superposition 97,125] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
287. ~r2_hidden(X13,X14) | v1_xtuple_0(X13) | ~v1_relat_1(X14) [superposition 104,97] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
290. 25 <=> ! [X0] : (~r2_hidden(sK2,X0) | ~v1_relat_1(X0)) [avatar definition]
291. ~r2_hidden(sK2,X0) | ~v1_relat_1(X0) <- (25) [avatar component clause 290] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
293. sK2 = k4_tarski(k1_xtuple_0(sK0),k2_xtuple_0(sK0)) | ~r2_hidden(sK2,X0) | ~v1_relat_1(X0) <- (2, 3) [forward demodulation 282,130] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
294. 25 | 24 | ~2 | ~3 [avatar split clause 293,128,123,264,290]
296. ~v1_relat_1(sK1) <- (5, 25) [resolution 291,140] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
298. $false <- (5, 6, 25) [subsumption resolution 296,145] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
299. ~5 | ~6 | ~25 [avatar contradiction clause 298]
303. sK0 = sK2 | ~v1_xtuple_0(sK0) <- (24) [superposition 91,266] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
313. 26 <=> v1_xtuple_0(sK0) [avatar definition]
315. ~v1_xtuple_0(sK0) <- (~26) [avatar component clause 313] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
318. ~v1_xtuple_0(sK0) <- (~1, 24) [subsumption resolution 303,120] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
319. ~26 | 1 | ~24 [avatar split clause 318,264,118,313]
338. v1_xtuple_0(sK0) | ~v1_relat_1(sK1) <- (4) [resolution 287,135] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
342. ~v1_relat_1(sK1) <- (4, ~26) [subsumption resolution 338,315] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
343. $false <- (4, 6, ~26) [subsumption resolution 342,145] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
344. ~4 | ~6 | 26 [avatar contradiction clause 343]
345. $false [avatar sat refutation 121,126,131,136,141,146,294,299,319,344] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:203,thDist:-203}
% SZS output end Proof for mcart_1__t89_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 37
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 146
% Active clauses: 52
% Passive clauses: 72
% Final active clauses: 49
% Final passive clauses: 19
% Main loop iterations started: 53
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 10
% Fw demodulations: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 13
% Fw demodulations to eq. taut.: 4
% 
% >>> Generating Inferences
% Binary resolution: 18
% Unit resulting resolution: 2
% Forward superposition: 19
% Backward superposition: 23
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 39
% SAT solver unit clauses: 15
% SAT solver binary clauses: 12
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

