i: [1,92,0,2,1,1,0,0] t1_grcat_1
a: [3,155,0,2,1,92]
s: 155
i: [1,91,0,2,1,1,0,0] t1_grcat_1
a: [3,160,0,2,1,91]
s: 160
i: [1,96,0,3,1,1,0,0] t1_grcat_1
a: [3,135,0,3,1,96]
s: 135
i: [1,95,0,3,1,1,0,0] t1_grcat_1
a: [3,140,0,3,1,95]
s: 140
i: [1,94,0,3,1,1,0,0] t1_grcat_1
a: [3,145,0,3,1,94]
s: 145
i: [1,93,0,3,1,1,0,0] t1_grcat_1
a: [3,150,0,3,1,93]
s: 150
i: [1,97,0,13,2,1,0,0] t1_grcat_1
a: [3,126,0,6,1,97]
s: 126
i: [1,98,0,2,1,0,0,255] fc1_xboole_0
a: [3,165,0,2,1,98]
s: 165
i: [1,119,0,2,1,0,0,1] rc2_xboole_0
a: [3,170,0,2,1,119]
s: 170
i: [1,120,0,2,1,0,0,1] rc1_xboole_0
a: [3,175,0,2,1,120]
s: 175
i: [1,122,0,2,1,0,0,1] rc1_classes2
a: [3,180,0,2,1,122]
s: 180
i: [1,121,0,2,1,0,0,1] rc1_classes2
a: [3,185,0,2,1,121]
s: 185
i: [1,99,0,4,2,0,0,255] cc2_classes2
s: 99
d: [2,187,1,2,1,1,40,99,155]
a: [3,192,1,2,1,187]
s: 192
i: [1,100,0,4,2,0,0,255] cc2_classes2
s: 100
i: [1,113,0,4,1,0,0,1] existence_m1_subset_1
s: 113
d: [2,199,1,2,1,1,40,100,155]
a: [3,204,1,2,1,199]
s: 204
i: [1,114,0,4,1,0,0,255] fc3_xboole_0
s: 114
d: [2,188,1,2,1,1,40,99,180]
a: [3,197,1,2,1,188]
s: 197
i: [1,101,0,5,2,0,0,255] t6_boole
s: 101
i: [1,108,0,5,2,0,0,1] t7_boole
s: 108
d: [2,200,1,2,1,1,40,100,180]
a: [3,209,1,2,1,200]
s: 209
i: [1,104,0,6,3,0,0,255] cc3_classes2
s: 104
i: [1,105,0,6,2,0,0,1] t1_subset
s: 105
i: [1,116,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 116
i: [1,103,0,7,3,0,0,1] rc2_classes2
s: 103
i: [1,107,0,7,3,0,0,2] t8_boole
s: 107
i: [1,115,0,7,1,0,0,255] commutativity_k2_tarski
s: 115
d: [2,211,1,7,3,0,44,101,101]
s: 211
i: [1,102,0,8,3,0,0,1] rc2_classes2
s: 102
i: [1,106,0,8,3,0,0,1] t2_subset
s: 106
d: [2,217,1,5,2,1,40,106,150]
d: [2,224,1,3,1,2,21,217,160]
a: [3,228,1,3,1,224]
s: 228
i: [1,109,0,10,1,0,0,1] d4_xtuple_0
s: 109
d: [2,218,1,5,2,1,40,106,145]
d: [2,230,1,3,1,2,21,218,160]
a: [3,234,1,3,1,230]
s: 234
i: [1,112,0,12,1,0,0,1] d8_xtuple_0
s: 112
i: [1,110,0,15,5,0,0,255] t58_classes2
s: 110
d: [2,219,1,5,2,1,40,106,140]
d: [2,236,1,3,1,2,21,219,160]
a: [3,240,1,3,1,236]
s: 240
d: [2,266,0,13,4,0,21,110,108]
s: 266
d: [2,220,1,5,2,1,40,106,135]
d: [2,242,1,3,1,2,21,220,160]
a: [3,246,1,3,1,242]
s: 246
i: [1,111,0,15,5,0,0,255] t58_classes2
s: 111
d: [2,288,0,13,4,0,21,111,108]
s: 288
d: [2,248,2,3,1,1,40,228,116]
a: [3,254,2,3,1,248]
s: 254
i: [1,117,0,16,5,0,0,255] dt_k6_classes2
s: 117
d: [2,257,2,3,1,1,40,234,116]
a: [3,263,2,3,1,257]
s: 263
i: [1,118,0,18,5,0,0,255] redefinition_k6_classes2
s: 118
d: [2,221,1,6,2,0,40,106,113]
s: 221
d: [2,267,2,3,1,1,40,240,116]
a: [3,273,2,3,1,267]
s: 273
d: [2,222,1,10,4,0,40,106,102]
d: [2,223,1,8,3,0,18,222]
s: 223
d: [2,280,2,3,1,1,40,246,116]
a: [3,286,2,3,1,280]
s: 286
d: [2,256,1,13,1,0,44,109,109]
s: 256
d: [2,304,1,12,1,0,22,256,112]
s: 304
d: [2,298,2,6,2,0,40,221,116]
s: 298
d: [2,275,1,13,4,0,40,266,116]
s: 275
d: [2,301,2,8,3,0,40,223,116]
s: 301
d: [2,276,1,13,4,0,40,266,105]
s: 276
d: [2,289,1,13,4,0,40,288,116]
s: 289
d: [2,290,1,13,4,0,40,288,105]
s: 290
d: [2,265,1,15,1,0,44,109,112]
s: 265
d: [2,296,1,25,9,0,44,117,118]
d: [2,297,1,15,5,0,18,296]
s: 297
d: [2,292,1,16,4,0,44,288,109]
s: 292
d: [2,294,1,18,6,0,40,117,106]
d: [2,295,1,16,5,0,18,294]
s: 295
d: [2,343,2,10,3,1,40,292,126]
d: [2,353,2,7,2,2,21,343,240]
d: [2,354,2,5,1,3,21,353,155]
a: [3,358,2,5,1,354]
s: 358
d: [2,370,3,0,0,4,52,155,228,234,358,288]
e: 370
d: [2,371,3,8,3,1,40,358,288]
d: [2,373,3,5,2,2,21,371,234]
d: [2,374,3,2,1,3,21,373,228]
d: [2,375,3,0,0,4,21,374,155]
e: 375
r: 126
r: 358
a: [3,125,0,6,1,97]
s: 125
a: [3,357,2,5,1,354]
s: 357
a: [3,130,0,7,1,97]
s: 130
d: [2,377,1,6,1,1,40,125,116]
a: [3,383,1,6,1,377]
s: 383
d: [2,390,3,5,1,1,40,357,116]
a: [3,396,3,5,1,390]
s: 396
d: [2,378,1,6,1,1,40,125,105]
a: [3,388,1,6,1,378]
s: 388
d: [2,391,3,5,1,1,40,357,105]
a: [3,401,3,5,1,391]
s: 401
d: [2,293,1,18,4,0,44,288,112]
s: 293
d: [2,405,2,0,0,4,52,155,246,125,130,293]
e: 405
d: [2,406,2,11,3,1,40,293,130]
d: [2,415,2,8,2,2,21,406,246]
d: [2,416,2,2,1,3,21,415,125]
d: [2,417,2,0,0,4,21,416,155]
e: 417
f: 417,375,354,242,236,230,224,91,92,93,94,95,96,97
% Refutation found. Thanks to Tanya!
% SZS status Theorem for grcat_1__t1_grcat_1
% SZS output start Proof for grcat_1__t1_grcat_1
1. ! [X0] : ((v1_classes2(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_subset_1(X2,X0) => ! [X3] : (m1_subset_1(X3,X0) => ! [X4] : (m1_subset_1(X4,X0) => (r2_hidden(k6_xtuple_0(X1,X2,X3,X4),X0) & r2_hidden(k3_xtuple_0(X1,X2,X3),X0))))))) [input t1_grcat_1]
2. ~! [X0] : ((v1_classes2(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_subset_1(X2,X0) => ! [X3] : (m1_subset_1(X3,X0) => ! [X4] : (m1_subset_1(X4,X0) => (r2_hidden(k6_xtuple_0(X1,X2,X3,X4),X0) & r2_hidden(k3_xtuple_0(X1,X2,X3),X0))))))) [negated conjecture 1]
10. ! [X0] : ! [X1] : ! [X2] : k3_xtuple_0(X0,X1,X2) = k4_tarski(k4_tarski(X0,X1),X2) [input d4_xtuple_0]
11. ! [X0] : ! [X1] : ! [X2] : ! [X3] : k6_xtuple_0(X0,X1,X2,X3) = k4_tarski(k3_xtuple_0(X0,X1,X2),X3) [input d8_xtuple_0]
34. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
35. ! [X0] : ! [X1] : ! [X2] : ((v1_classes2(X2) & ~v1_xboole_0(X2)) => ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) => (r2_hidden(k4_tarski(X0,X1),X2) & r2_hidden(k2_tarski(X0,X1),X2)))) [input t58_classes2]
37. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
40. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 34]
42. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 37]
43. ! [X0,X1,X2] : k3_xtuple_0(X0,X1,X2) = k4_tarski(k4_tarski(X0,X1),X2) [flattening 10]
44. ! [X0,X1,X2] : ((v1_classes2(X2) & ~v1_xboole_0(X2)) => ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) => (r2_hidden(k4_tarski(X0,X1),X2) & r2_hidden(k2_tarski(X0,X1),X2)))) [flattening 35]
45. ! [X0,X1,X2,X3] : k6_xtuple_0(X0,X1,X2,X3) = k4_tarski(k3_xtuple_0(X0,X1,X2),X3) [flattening 11]
55. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : ((~r2_hidden(k6_xtuple_0(X1,X2,X3,X4),X0) | ~r2_hidden(k3_xtuple_0(X1,X2,X3),X0)) & m1_subset_1(X4,X0)) & m1_subset_1(X3,X0)) & m1_subset_1(X2,X0)) & m1_subset_1(X1,X0)) & (v1_classes2(X0) & ~v1_xboole_0(X0))) [ennf transformation 2]
56. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : ((~r2_hidden(k6_xtuple_0(X1,X2,X3,X4),X0) | ~r2_hidden(k3_xtuple_0(X1,X2,X3),X0)) & m1_subset_1(X4,X0)) & m1_subset_1(X3,X0)) & m1_subset_1(X2,X0)) & m1_subset_1(X1,X0)) & v1_classes2(X0) & ~v1_xboole_0(X0)) [flattening 55]
64. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 40]
65. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 64]
67. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 42]
68. ! [X0,X1,X2] : (((r2_hidden(k4_tarski(X0,X1),X2) & r2_hidden(k2_tarski(X0,X1),X2)) | (~r2_hidden(X1,X2) | ~r2_hidden(X0,X2))) | (~v1_classes2(X2) | v1_xboole_0(X2))) [ennf transformation 44]
69. ! [X0,X1,X2] : ((r2_hidden(k4_tarski(X0,X1),X2) & r2_hidden(k2_tarski(X0,X1),X2)) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~v1_classes2(X2) | v1_xboole_0(X2)) [flattening 68]
75. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : ((~r2_hidden(k6_xtuple_0(X1,X2,X3,X4),X0) | ~r2_hidden(k3_xtuple_0(X1,X2,X3),X0)) & m1_subset_1(X4,X0)) & m1_subset_1(X3,X0)) & m1_subset_1(X2,X0)) & m1_subset_1(X1,X0)) & v1_classes2(X0) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : ((~r2_hidden(k6_xtuple_0(X1,X2,X3,X4),sK0) | ~r2_hidden(k3_xtuple_0(X1,X2,X3),sK0)) & m1_subset_1(X4,sK0)) & m1_subset_1(X3,sK0)) & m1_subset_1(X2,sK0)) & m1_subset_1(X1,sK0)) & v1_classes2(sK0) & ~v1_xboole_0(sK0)) [choice axiom]
76. ? [X1] : (? [X2] : (? [X3] : (? [X4] : ((~r2_hidden(k6_xtuple_0(X1,X2,X3,X4),sK0) | ~r2_hidden(k3_xtuple_0(X1,X2,X3),sK0)) & m1_subset_1(X4,sK0)) & m1_subset_1(X3,sK0)) & m1_subset_1(X2,sK0)) & m1_subset_1(X1,sK0)) => (? [X2] : (? [X3] : (? [X4] : ((~r2_hidden(k6_xtuple_0(sK1,X2,X3,X4),sK0) | ~r2_hidden(k3_xtuple_0(sK1,X2,X3),sK0)) & m1_subset_1(X4,sK0)) & m1_subset_1(X3,sK0)) & m1_subset_1(X2,sK0)) & m1_subset_1(sK1,sK0)) [choice axiom]
77. ? [X2] : (? [X3] : (? [X4] : ((~r2_hidden(k6_xtuple_0(sK1,X2,X3,X4),sK0) | ~r2_hidden(k3_xtuple_0(sK1,X2,X3),sK0)) & m1_subset_1(X4,sK0)) & m1_subset_1(X3,sK0)) & m1_subset_1(X2,sK0)) => (? [X3] : (? [X4] : ((~r2_hidden(k6_xtuple_0(sK1,sK2,X3,X4),sK0) | ~r2_hidden(k3_xtuple_0(sK1,sK2,X3),sK0)) & m1_subset_1(X4,sK0)) & m1_subset_1(X3,sK0)) & m1_subset_1(sK2,sK0)) [choice axiom]
78. ? [X3] : (? [X4] : ((~r2_hidden(k6_xtuple_0(sK1,sK2,X3,X4),sK0) | ~r2_hidden(k3_xtuple_0(sK1,sK2,X3),sK0)) & m1_subset_1(X4,sK0)) & m1_subset_1(X3,sK0)) => (? [X4] : ((~r2_hidden(k6_xtuple_0(sK1,sK2,sK3,X4),sK0) | ~r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0)) & m1_subset_1(X4,sK0)) & m1_subset_1(sK3,sK0)) [choice axiom]
79. ? [X4] : ((~r2_hidden(k6_xtuple_0(sK1,sK2,sK3,X4),sK0) | ~r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0)) & m1_subset_1(X4,sK0)) => ((~r2_hidden(k6_xtuple_0(sK1,sK2,sK3,sK4),sK0) | ~r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0)) & m1_subset_1(sK4,sK0)) [choice axiom]
80. (((((~r2_hidden(k6_xtuple_0(sK1,sK2,sK3,sK4),sK0) | ~r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0)) & m1_subset_1(sK4,sK0)) & m1_subset_1(sK3,sK0)) & m1_subset_1(sK2,sK0)) & m1_subset_1(sK1,sK0)) & v1_classes2(sK0) & ~v1_xboole_0(sK0) [skolemisation 56,79,78,77,76,75]
91. ~v1_xboole_0(sK0) [cnf transformation 80] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
92. v1_classes2(sK0) [cnf transformation 80] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
93. m1_subset_1(sK1,sK0) [cnf transformation 80] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
94. m1_subset_1(sK2,sK0) [cnf transformation 80] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
95. m1_subset_1(sK3,sK0) [cnf transformation 80] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
96. m1_subset_1(sK4,sK0) [cnf transformation 80] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
97. ~r2_hidden(k6_xtuple_0(sK1,sK2,sK3,sK4),sK0) | ~r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0) [cnf transformation 80] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
106. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 65] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
108. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 67] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
109. k3_xtuple_0(X0,X1,X2) = k4_tarski(k4_tarski(X0,X1),X2) [cnf transformation 43] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
111. r2_hidden(k4_tarski(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~v1_classes2(X2) | v1_xboole_0(X2) [cnf transformation 69] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:1,thDist:-1}
112. k6_xtuple_0(X0,X1,X2,X3) = k4_tarski(k3_xtuple_0(X0,X1,X2),X3) [cnf transformation 45] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
124. 1 <=> r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0) [avatar definition]
125. r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0) <- (1) [avatar component clause 124] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
126. ~r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0) <- (~1) [avatar component clause 124] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
128. 2 <=> r2_hidden(k6_xtuple_0(sK1,sK2,sK3,sK4),sK0) [avatar definition]
130. ~r2_hidden(k6_xtuple_0(sK1,sK2,sK3,sK4),sK0) <- (~2) [avatar component clause 128] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
131. ~1 | ~2 [avatar split clause 97,128,124]
133. 3 <=> m1_subset_1(sK4,sK0) [avatar definition]
135. m1_subset_1(sK4,sK0) <- (3) [avatar component clause 133] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
136. 3 [avatar split clause 96,133]
138. 4 <=> m1_subset_1(sK3,sK0) [avatar definition]
140. m1_subset_1(sK3,sK0) <- (4) [avatar component clause 138] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
141. 4 [avatar split clause 95,138]
143. 5 <=> m1_subset_1(sK2,sK0) [avatar definition]
145. m1_subset_1(sK2,sK0) <- (5) [avatar component clause 143] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
146. 5 [avatar split clause 94,143]
148. 6 <=> m1_subset_1(sK1,sK0) [avatar definition]
150. m1_subset_1(sK1,sK0) <- (6) [avatar component clause 148] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
151. 6 [avatar split clause 93,148]
153. 7 <=> v1_classes2(sK0) [avatar definition]
155. v1_classes2(sK0) <- (7) [avatar component clause 153] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
156. 7 [avatar split clause 92,153]
158. 8 <=> v1_xboole_0(sK0) [avatar definition]
160. ~v1_xboole_0(sK0) <- (~8) [avatar component clause 158] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
161. ~8 [avatar split clause 91,158]
217. v1_xboole_0(sK0) | r2_hidden(sK1,sK0) <- (6) [resolution 106,150] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
218. v1_xboole_0(sK0) | r2_hidden(sK2,sK0) <- (5) [resolution 106,145] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
219. v1_xboole_0(sK0) | r2_hidden(sK3,sK0) <- (4) [resolution 106,140] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
220. v1_xboole_0(sK0) | r2_hidden(sK4,sK0) <- (3) [resolution 106,135] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
224. r2_hidden(sK1,sK0) <- (6, ~8) [subsumption resolution 217,160] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
226. 18 <=> r2_hidden(sK1,sK0) [avatar definition]
228. r2_hidden(sK1,sK0) <- (18) [avatar component clause 226] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
229. 18 | ~6 | 8 [avatar split clause 224,158,148,226]
230. r2_hidden(sK2,sK0) <- (5, ~8) [subsumption resolution 218,160] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
232. 19 <=> r2_hidden(sK2,sK0) [avatar definition]
234. r2_hidden(sK2,sK0) <- (19) [avatar component clause 232] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
235. 19 | ~5 | 8 [avatar split clause 230,158,143,232]
236. r2_hidden(sK3,sK0) <- (4, ~8) [subsumption resolution 219,160] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
238. 20 <=> r2_hidden(sK3,sK0) [avatar definition]
240. r2_hidden(sK3,sK0) <- (20) [avatar component clause 238] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
241. 20 | ~4 | 8 [avatar split clause 236,158,138,238]
242. r2_hidden(sK4,sK0) <- (3, ~8) [subsumption resolution 220,160] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
244. 21 <=> r2_hidden(sK4,sK0) [avatar definition]
246. r2_hidden(sK4,sK0) <- (21) [avatar component clause 244] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
247. 21 | ~3 | 8 [avatar split clause 242,158,133,244]
288. r2_hidden(k4_tarski(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~v1_classes2(X2) [subsumption resolution 111,108] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
292. r2_hidden(k3_xtuple_0(X0,X1,X2),X3) | ~r2_hidden(X2,X3) | ~r2_hidden(k4_tarski(X0,X1),X3) | ~v1_classes2(X3) [superposition 288,109] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
293. r2_hidden(k6_xtuple_0(X4,X5,X6,X7),X8) | ~r2_hidden(X7,X8) | ~r2_hidden(k3_xtuple_0(X4,X5,X6),X8) | ~v1_classes2(X8) [superposition 288,112] {a:1,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
343. ~r2_hidden(sK3,sK0) | ~r2_hidden(k4_tarski(sK1,sK2),sK0) | ~v1_classes2(sK0) <- (~1) [resolution 292,126] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
353. ~r2_hidden(k4_tarski(sK1,sK2),sK0) | ~v1_classes2(sK0) <- (~1, 20) [subsumption resolution 343,240] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
354. ~r2_hidden(k4_tarski(sK1,sK2),sK0) <- (~1, 7, 20) [subsumption resolution 353,155] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
356. 26 <=> r2_hidden(k4_tarski(sK1,sK2),sK0) [avatar definition]
358. ~r2_hidden(k4_tarski(sK1,sK2),sK0) <- (~26) [avatar component clause 356] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
359. ~26 | 1 | ~7 | ~20 [avatar split clause 354,238,153,124,356]
371. ~r2_hidden(sK2,sK0) | ~r2_hidden(sK1,sK0) | ~v1_classes2(sK0) <- (~26) [resolution 358,288] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
373. ~r2_hidden(sK1,sK0) | ~v1_classes2(sK0) <- (19, ~26) [subsumption resolution 371,234] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
374. ~v1_classes2(sK0) <- (18, 19, ~26) [subsumption resolution 373,228] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
375. $false <- (7, 18, 19, ~26) [subsumption resolution 374,155] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
376. ~7 | ~18 | ~19 | 26 [avatar contradiction clause 375]
406. ~r2_hidden(sK4,sK0) | ~r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0) | ~v1_classes2(sK0) <- (~2) [resolution 293,130] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
415. ~r2_hidden(k3_xtuple_0(sK1,sK2,sK3),sK0) | ~v1_classes2(sK0) <- (~2, 21) [subsumption resolution 406,246] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
416. ~v1_classes2(sK0) <- (1, ~2, 21) [subsumption resolution 415,125] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
417. $false <- (1, ~2, 7, 21) [subsumption resolution 416,155] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
418. ~1 | 2 | ~7 | ~21 [avatar contradiction clause 417]
419. $false [avatar sat refutation 131,136,141,146,151,156,161,229,235,241,247,359,376,418] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:297,thDist:-297}
% SZS output end Proof for grcat_1__t1_grcat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 37
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 32
% Generated clauses: 207
% Active clauses: 66
% Passive clauses: 106
% Final active clauses: 64
% Final passive clauses: 36
% Main loop iterations started: 70
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 19
% Fw demodulations: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 33
% 
% >>> Generating Inferences
% Binary resolution: 82
% Unit resulting resolution: 2
% Forward superposition: 20
% Backward superposition: 7
% Self superposition: 3
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 35
% SAT solver unit clauses: 11
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 5117
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% congruence closure: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.005 s ( own 0.001 s ) 

