i: [1,61,0,2,1,1,0,0] t97_funct_4
a: [3,98,0,2,1,61]
s: 98
i: [1,60,0,2,1,1,0,0] t97_funct_4
a: [3,103,0,2,1,60]
s: 103
i: [1,59,0,2,1,1,0,0] t97_funct_4
a: [3,108,0,2,1,59]
s: 108
i: [1,58,0,2,1,1,0,0] t97_funct_4
a: [3,113,0,2,1,58]
s: 113
i: [1,62,0,3,1,1,0,0] t97_funct_4
a: [3,93,0,3,1,62]
s: 93
i: [1,63,0,5,1,1,0,0] t97_funct_4
a: [3,88,0,5,1,63]
s: 88
i: [1,84,0,2,1,0,0,1] rc1_funct_1
a: [3,129,0,2,1,84]
s: 129
i: [1,83,0,2,1,0,0,1] rc1_funct_1
a: [3,134,0,2,1,83]
s: 134
i: [1,73,0,3,1,0,0,1] reflexivity_r1_tarski
s: 73
i: [1,72,0,4,1,0,0,255] existence_m1_subset_1
s: 72
i: [1,74,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 74
i: [1,70,0,7,2,0,0,255] t3_subset
s: 70
i: [1,71,0,7,2,0,0,255] t3_subset
s: 71
d: [2,136,1,5,1,0,40,70,72]
s: 136
i: [1,75,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 75
i: [1,76,0,11,5,0,0,255] reflexivity_r1_partfun1
a: [3,120,0,7,3,76]
s: 120
i: [1,64,0,8,3,0,0,255] cc2_relat_1
s: 64
i: [1,69,0,8,2,0,0,1] t12_xboole_1
s: 69
d: [2,138,1,6,2,0,40,64,72]
s: 138
i: [1,82,0,8,3,0,0,2] fc3_relat_1
s: 82
i: [1,77,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,124,0,9,3,77]
s: 124
d: [2,139,1,7,3,0,40,64,71]
s: 139
i: [1,65,0,10,4,0,0,255] cc3_funct_1
s: 65
d: [2,140,1,8,2,0,44,69,75]
s: 140
i: [1,79,0,12,5,0,0,1] dt_k1_funct_4
s: 79
i: [1,80,0,12,5,0,0,1] dt_k1_funct_4
s: 80
d: [2,141,1,8,2,0,44,69,75]
s: 141
i: [1,66,0,14,6,0,0,255] t54_partfun1
s: 66
d: [2,142,1,8,2,0,44,75,69]
s: 142
d: [2,165,0,12,5,0,21,66,139]
s: 165
i: [1,81,0,14,6,0,0,255] symmetry_r1_partfun1
s: 81
d: [2,166,1,9,4,1,40,165,93]
d: [2,170,1,7,3,2,21,166,98]
d: [2,171,1,5,2,3,21,170,103]
d: [2,172,1,3,1,4,21,171,108]
a: [3,176,1,3,1,172]
s: 176
i: [1,78,0,17,5,0,0,1] rd3_funct_4
s: 78
d: [2,180,2,11,5,1,40,176,81]
d: [2,181,2,9,4,2,21,180,98]
d: [2,182,2,7,3,3,21,181,103]
d: [2,183,2,5,2,4,21,182,108]
d: [2,184,2,3,1,5,21,183,113]
a: [3,188,2,3,1,184]
s: 188
i: [1,67,0,18,6,0,0,255] t30_funct_4
s: 67
i: [1,68,0,18,6,0,0,255] t30_funct_4
s: 68
d: [2,143,1,8,2,0,44,75,69]
s: 143
d: [2,152,1,8,3,0,40,65,72]
s: 152
d: [2,153,1,9,4,0,40,65,71]
s: 153
d: [2,201,1,19,7,0,44,67,69]
d: [2,214,1,17,6,0,21,201,139]
d: [2,215,1,14,5,0,21,214,165]
s: 215
d: [2,231,1,12,4,0,21,215,153]
s: 231
d: [2,232,2,0,0,4,52,98,103,93,88,231]
e: 232
d: [2,234,2,10,4,1,44,88,231]
d: [2,243,2,7,3,1,19,234]
d: [2,246,2,5,2,2,21,243,98]
d: [2,247,2,3,1,3,21,246,103]
d: [2,248,2,0,0,4,21,247,93]
e: 248
f: 248,60,61,62,63
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t97_funct_4
% SZS output start Proof for funct_4__t97_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) => k1_funct_4(X0,X1) = X1))) [input t97_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) => k1_funct_4(X0,X1) = X1))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
4. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_funct_1(X1))) [input cc3_funct_1]
19. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k2_xboole_0(X0,X1) = X1) [input t12_xboole_1]
20. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) <=> k1_funct_4(X0,X1) = k2_xboole_0(X0,X1)))) [input t30_funct_4]
21. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
22. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) => r1_partfun1(X0,X1)))) [input t54_partfun1]
23. ! [X0,X1] : (r1_tarski(X0,X1) => k2_xboole_0(X0,X1) = X1) [flattening 19]
24. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 21]
27. ? [X0] : (? [X1] : ((k1_funct_4(X0,X1) != X1 & r1_tarski(X0,X1)) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
28. ? [X0] : (? [X1] : (k1_funct_4(X0,X1) != X1 & r1_tarski(X0,X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 27]
29. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
30. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 4]
31. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 30]
32. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) | ~r1_tarski(X0,X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 22]
33. ! [X0] : (! [X1] : (r1_partfun1(X0,X1) | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 32]
34. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> k1_funct_4(X0,X1) = k2_xboole_0(X0,X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 20]
35. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> k1_funct_4(X0,X1) = k2_xboole_0(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 34]
36. ! [X0,X1] : (k2_xboole_0(X0,X1) = X1 | ~r1_tarski(X0,X1)) [ennf transformation 23]
49. ? [X0] : (? [X1] : (k1_funct_4(X0,X1) != X1 & r1_tarski(X0,X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (k1_funct_4(sK0,X1) != X1 & r1_tarski(sK0,X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
50. ? [X1] : (k1_funct_4(sK0,X1) != X1 & r1_tarski(sK0,X1) & v1_funct_1(X1) & v1_relat_1(X1)) => (sK1 != k1_funct_4(sK0,sK1) & r1_tarski(sK0,sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
51. (sK1 != k1_funct_4(sK0,sK1) & r1_tarski(sK0,sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 28,50,49]
52. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | k1_funct_4(X0,X1) != k2_xboole_0(X0,X1)) & (k1_funct_4(X0,X1) = k2_xboole_0(X0,X1) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 35]
53. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 24]
60. v1_relat_1(sK1) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
61. v1_funct_1(sK1) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
62. r1_tarski(sK0,sK1) [cnf transformation 51] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
63. sK1 != k1_funct_4(sK0,sK1) [cnf transformation 51] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
64. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
65. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_funct_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
66. r1_partfun1(X0,X1) | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
67. k1_funct_4(X0,X1) = k2_xboole_0(X0,X1) | ~r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 52] {a:0,w:18,wCS:7200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
69. k2_xboole_0(X0,X1) = X1 | ~r1_tarski(X0,X1) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 53] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
86. 1 <=> sK1 = k1_funct_4(sK0,sK1) [avatar definition]
88. sK1 != k1_funct_4(sK0,sK1) <- (~1) [avatar component clause 86] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. ~1 [avatar split clause 63,86]
91. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
93. r1_tarski(sK0,sK1) <- (2) [avatar component clause 91] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
94. 2 [avatar split clause 62,91]
96. 3 <=> v1_funct_1(sK1) [avatar definition]
98. v1_funct_1(sK1) <- (3) [avatar component clause 96] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
99. 3 [avatar split clause 61,96]
101. 4 <=> v1_relat_1(sK1) [avatar definition]
103. v1_relat_1(sK1) <- (4) [avatar component clause 101] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. 4 [avatar split clause 60,101]
139. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 64,71] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
153. ~r1_tarski(X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | v1_funct_1(X1) [resolution 65,71] {a:1,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
165. ~r1_tarski(X0,X1) | r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) [subsumption resolution 66,139] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
201. k1_funct_4(X5,X6) = X6 | ~r1_partfun1(X5,X6) | ~v1_funct_1(X6) | ~v1_relat_1(X6) | ~v1_funct_1(X5) | ~v1_relat_1(X5) | ~r1_tarski(X5,X6) [superposition 67,69] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:2,thDist:-2}
214. k1_funct_4(X5,X6) = X6 | ~r1_partfun1(X5,X6) | ~v1_funct_1(X6) | ~v1_relat_1(X6) | ~v1_funct_1(X5) | ~r1_tarski(X5,X6) [subsumption resolution 201,139] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
215. k1_funct_4(X5,X6) = X6 | ~v1_funct_1(X6) | ~v1_relat_1(X6) | ~v1_funct_1(X5) | ~r1_tarski(X5,X6) [subsumption resolution 214,165] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
231. k1_funct_4(X5,X6) = X6 | ~v1_funct_1(X6) | ~v1_relat_1(X6) | ~r1_tarski(X5,X6) [subsumption resolution 215,153] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
234. sK1 != sK1 | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~r1_tarski(sK0,sK1) <- (~1) [superposition 88,231] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
243. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~r1_tarski(sK0,sK1) <- (~1) [trivial inequality removal 234] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
246. ~v1_relat_1(sK1) | ~r1_tarski(sK0,sK1) <- (~1, 3) [subsumption resolution 243,98] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
247. ~r1_tarski(sK0,sK1) <- (~1, 3, 4) [subsumption resolution 246,103] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
248. $false <- (~1, 2, 3, 4) [subsumption resolution 247,93] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
249. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 248]
250. $false [avatar sat refutation 89,94,99,104,249] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:75,thDist:-75}
% SZS output end Proof for funct_4__t97_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 141
% Active clauses: 36
% Passive clauses: 54
% Final active clauses: 36
% Final passive clauses: 13
% Main loop iterations started: 41
% 
% >>> Simplifying Inferences
% Duplicate literals: 46
% Trivial inequalities: 1
% Fw subsumption resolutions: 19
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 24
% 
% >>> Generating Inferences
% Binary resolution: 18
% Unit resulting resolution: 1
% Forward superposition: 24
% Backward superposition: 19
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 8
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

