i: [1,47,0,3,1,1,0,0] t22_zfmisc_1
a: [3,73,0,3,1,47]
s: 73
i: [1,46,0,3,1,1,0,0] t22_zfmisc_1
a: [3,78,0,3,1,46]
s: 78
i: [1,45,0,7,1,1,0,0] t22_zfmisc_1
a: [3,83,0,7,1,45]
s: 83
i: [1,48,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,88,0,2,1,48]
s: 88
i: [1,49,0,2,1,0,0,2] fc1_xboole_0
a: [3,93,0,2,1,49]
s: 93
i: [1,64,0,2,1,0,0,2] rc2_xboole_0
a: [3,103,0,2,1,64]
s: 103
i: [1,65,0,2,1,0,0,2] rc1_xboole_0
a: [3,108,0,2,1,65]
s: 108
i: [1,51,0,3,1,0,0,2] fc2_xboole_0
s: 51
i: [1,61,0,3,1,0,0,1] reflexivity_r1_tarski
s: 61
i: [1,50,0,3,1,0,0,255] d2_xboole_0
a: [3,98,0,3,1,50]
s: 98
r: 88
i: [1,62,0,4,1,0,0,1] fc3_xboole_0
s: 62
i: [1,52,0,5,2,0,0,2] t6_boole
s: 52
i: [1,69,0,5,1,0,0,1] l38_zfmisc_1
s: 69
i: [1,67,0,6,1,0,0,1] l38_zfmisc_1
s: 67
i: [1,68,0,6,1,0,0,1] l38_zfmisc_1
s: 68
i: [1,53,0,7,3,0,0,1] t8_boole
s: 53
i: [1,63,0,7,1,0,0,1] commutativity_k2_tarski
s: 63
i: [1,66,0,7,1,0,0,1] l38_zfmisc_1
s: 66
d: [2,111,1,7,3,0,44,52,52]
s: 111
i: [1,54,0,9,2,0,0,2] t4_zfmisc_1
s: 54
i: [1,60,0,13,3,0,0,1] t6_zfmisc_1
s: 60
d: [2,113,1,7,2,0,44,69,52]
s: 113
i: [1,55,0,21,5,0,0,1] l38_zfmisc_1
s: 55
d: [2,133,1,24,4,1,40,55,83]
a: [3,146,1,5,1,133]
s: 146
r: 83
d: [2,161,2,2,1,1,44,62,146]
d: [2,167,2,0,0,2,21,161,93]
e: 167
r: 146
s: 83
a: [3,145,1,5,1,133]
s: 145
a: [3,150,1,6,1,133]
s: 150
r: 83
d: [2,172,2,0,0,2,52,78,150,54]
e: 172
r: 150
s: 83
a: [3,149,1,6,1,133]
s: 149
a: [3,154,1,6,1,133]
s: 154
r: 83
d: [2,188,2,0,0,2,52,73,154,54]
e: 188
r: 154
s: 83
a: [3,153,1,6,1,133]
s: 153
a: [3,142,1,7,1,133]
s: 142
r: 83
d: [2,206,2,0,0,3,52,73,78,142,60]
e: 206
d: [2,201,1,24,4,1,40,83,55]
d: [2,202,1,18,3,2,21,201,149]
d: [2,203,1,13,2,3,21,202,145]
f: 206,203,188,172,167,49,45,46,47
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t22_zfmisc_1
% SZS output start Proof for zfmisc_1__t22_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ~(X0 != X3 & X0 != X2 & r1_tarski(k2_tarski(X0,X1),k2_tarski(X2,X3))) [input t22_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ~(X0 != X3 & X0 != X2 & r1_tarski(k2_tarski(X0,X1),k2_tarski(X2,X3))) [negated conjecture 1]
10. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
12. ! [X0,X1] : ~v1_xboole_0(k2_tarski(X0,X1)) [input fc3_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,k2_tarski(X1,X2)) <=> ~(k2_tarski(X1,X2) != X0 & k1_tarski(X2) != X0 & k1_tarski(X1) != X0 & k1_xboole_0 != X0)) [input l38_zfmisc_1]
17. ! [X0] : ! [X1] : ! [X2] : (k1_tarski(X0) = k2_tarski(X1,X2) => X0 = X1) [input t4_zfmisc_1]
19. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ~(X0 != X3 & X0 != X2 & k2_tarski(X0,X1) = k2_tarski(X2,X3)) [input t6_zfmisc_1]
22. ~! [X0,X1,X2,X3] : ~(X0 != X3 & X0 != X2 & r1_tarski(k2_tarski(X0,X1),k2_tarski(X2,X3))) [flattening 2]
25. ! [X0,X1,X2] : (k1_tarski(X0) = k2_tarski(X1,X2) => X0 = X1) [flattening 17]
26. ! [X0,X1,X2] : (r1_tarski(X0,k2_tarski(X1,X2)) <=> ~(k2_tarski(X1,X2) != X0 & k1_tarski(X2) != X0 & k1_tarski(X1) != X0 & k1_xboole_0 != X0)) [flattening 13]
27. ! [X0,X1,X2,X3] : ~(X0 != X3 & X0 != X2 & k2_tarski(X0,X1) = k2_tarski(X2,X3)) [flattening 19]
31. ? [X0,X1,X2,X3] : (X0 != X3 & X0 != X2 & r1_tarski(k2_tarski(X0,X1),k2_tarski(X2,X3))) [ennf transformation 22]
34. ! [X0,X1,X2] : (X0 = X1 | k1_tarski(X0) != k2_tarski(X1,X2)) [ennf transformation 25]
35. ! [X0,X1,X2] : (r1_tarski(X0,k2_tarski(X1,X2)) <=> (k2_tarski(X1,X2) = X0 | k1_tarski(X2) = X0 | k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [ennf transformation 26]
36. ! [X0,X1,X2,X3] : (X0 = X3 | X0 = X2 | k2_tarski(X0,X1) != k2_tarski(X2,X3)) [ennf transformation 27]
37. ? [X0,X1,X2,X3] : (X0 != X3 & X0 != X2 & r1_tarski(k2_tarski(X0,X1),k2_tarski(X2,X3))) => (sK0 != sK3 & sK0 != sK2 & r1_tarski(k2_tarski(sK0,sK1),k2_tarski(sK2,sK3))) [choice axiom]
38. sK0 != sK3 & sK0 != sK2 & r1_tarski(k2_tarski(sK0,sK1),k2_tarski(sK2,sK3)) [skolemisation 31,37]
39. ! [X0,X1,X2] : ((r1_tarski(X0,k2_tarski(X1,X2)) | (k2_tarski(X1,X2) != X0 & k1_tarski(X2) != X0 & k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & ((k2_tarski(X1,X2) = X0 | k1_tarski(X2) = X0 | k1_tarski(X1) = X0 | k1_xboole_0 = X0) | ~r1_tarski(X0,k2_tarski(X1,X2)))) [nnf transformation 35]
40. ! [X0,X1,X2] : ((r1_tarski(X0,k2_tarski(X1,X2)) | (k2_tarski(X1,X2) != X0 & k1_tarski(X2) != X0 & k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & (k2_tarski(X1,X2) = X0 | k1_tarski(X2) = X0 | k1_tarski(X1) = X0 | k1_xboole_0 = X0 | ~r1_tarski(X0,k2_tarski(X1,X2)))) [flattening 39]
45. r1_tarski(k2_tarski(sK0,sK1),k2_tarski(sK2,sK3)) [cnf transformation 38] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. sK0 != sK2 [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. sK0 != sK3 [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. v1_xboole_0(k1_xboole_0) [cnf transformation 10] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
54. k1_tarski(X0) != k2_tarski(X1,X2) | X0 = X1 [cnf transformation 34] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
55. ~r1_tarski(X0,k2_tarski(X1,X2)) | k1_tarski(X2) = X0 | k1_tarski(X1) = X0 | k1_xboole_0 = X0 | k2_tarski(X1,X2) = X0 [cnf transformation 40] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. k2_tarski(X0,X1) != k2_tarski(X2,X3) | X0 = X2 | X0 = X3 [cnf transformation 36] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. ~v1_xboole_0(k2_tarski(X0,X1)) [cnf transformation 12] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. 1 <=> sK0 = sK3 [avatar definition]
73. sK0 != sK3 <- (~1) [avatar component clause 71] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
74. ~1 [avatar split clause 47,71]
76. 2 <=> sK0 = sK2 [avatar definition]
78. sK0 != sK2 <- (~2) [avatar component clause 76] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. ~2 [avatar split clause 46,76]
81. 3 <=> r1_tarski(k2_tarski(sK0,sK1),k2_tarski(sK2,sK3)) [avatar definition]
83. r1_tarski(k2_tarski(sK0,sK1),k2_tarski(sK2,sK3)) <- (3) [avatar component clause 81] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. 3 [avatar split clause 45,81]
91. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
93. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 91] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
94. 5 [avatar split clause 49,91]
140. 9 <=> k2_tarski(sK0,sK1) = k2_tarski(sK2,sK3) [avatar definition]
142. k2_tarski(sK0,sK1) = k2_tarski(sK2,sK3) <- (9) [avatar component clause 140] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. 10 <=> k1_xboole_0 = k2_tarski(sK0,sK1) [avatar definition]
145. k1_xboole_0 != k2_tarski(sK0,sK1) <- (~10) [avatar component clause 144] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
146. k1_xboole_0 = k2_tarski(sK0,sK1) <- (10) [avatar component clause 144] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
148. 11 <=> k2_tarski(sK0,sK1) = k1_tarski(sK2) [avatar definition]
149. k2_tarski(sK0,sK1) != k1_tarski(sK2) <- (~11) [avatar component clause 148] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
150. k2_tarski(sK0,sK1) = k1_tarski(sK2) <- (11) [avatar component clause 148] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
152. 12 <=> k2_tarski(sK0,sK1) = k1_tarski(sK3) [avatar definition]
154. k2_tarski(sK0,sK1) = k1_tarski(sK3) <- (12) [avatar component clause 152] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
161. ~v1_xboole_0(k1_xboole_0) <- (10) [superposition 62,146] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
167. $false <- (5, 10) [subsumption resolution 161,93] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
168. ~5 | ~10 [avatar contradiction clause 167]
172. $false <- (~2, 11) [unit resulting resolution 78,150,54] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
183. 2 | ~11 [avatar contradiction clause 172]
188. $false <- (~1, 12) [unit resulting resolution 73,154,54] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
200. 1 | ~12 [avatar contradiction clause 188]
201. k2_tarski(sK0,sK1) = k1_tarski(sK3) | k2_tarski(sK0,sK1) = k1_tarski(sK2) | k1_xboole_0 = k2_tarski(sK0,sK1) | k2_tarski(sK0,sK1) = k2_tarski(sK2,sK3) <- (3) [resolution 83,55] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
202. k2_tarski(sK0,sK1) = k1_tarski(sK3) | k1_xboole_0 = k2_tarski(sK0,sK1) | k2_tarski(sK0,sK1) = k2_tarski(sK2,sK3) <- (3, ~11) [subsumption resolution 201,149] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
203. k2_tarski(sK0,sK1) = k1_tarski(sK3) | k2_tarski(sK0,sK1) = k2_tarski(sK2,sK3) <- (3, ~10, ~11) [subsumption resolution 202,145] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
204. 9 | 12 | ~3 | 10 | 11 [avatar split clause 203,148,144,81,152,140]
206. $false <- (~1, ~2, 9) [unit resulting resolution 73,78,142,60] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
216. 1 | 2 | ~9 [avatar contradiction clause 206]
217. $false [avatar sat refutation 74,79,84,94,168,183,200,204,216] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:77,thDist:-77}
% SZS output end Proof for zfmisc_1__t22_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 122
% Active clauses: 31
% Passive clauses: 41
% Final active clauses: 23
% Final passive clauses: 4
% Main loop iterations started: 33
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 4
% Bw demodulations: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Equational tautologies: 5
% Forward subsumptions: 28
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 3
% Forward superposition: 11
% Backward superposition: 48
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 13
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 8
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward demodulation: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

