i: [1,92,0,12,4,1,0,0] t52_mcart_1
a: [3,124,0,3,1,92]
s: 124
i: [1,91,0,3,1,1,0,0] t52_mcart_1
a: [3,141,0,3,1,91]
s: 141
i: [1,90,0,3,1,1,0,0] t52_mcart_1
a: [3,146,0,3,1,90]
s: 146
i: [1,89,0,3,1,1,0,0] t52_mcart_1
a: [3,151,0,3,1,89]
s: 151
i: [1,88,0,3,1,1,0,0] t52_mcart_1
a: [3,156,0,3,1,88]
s: 156
i: [1,93,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,166,0,2,1,93]
s: 166
i: [1,94,0,2,1,0,0,1] fc1_xboole_0
a: [3,171,0,2,1,94]
s: 171
i: [1,116,0,2,1,0,0,2] rc2_xboole_0
a: [3,181,0,2,1,116]
s: 181
i: [1,117,0,2,1,0,0,2] rc1_xboole_0
a: [3,186,0,2,1,117]
s: 186
i: [1,119,0,2,1,0,0,3] rc1_relat_1
a: [3,191,0,2,1,119]
s: 191
i: [1,118,0,2,1,0,0,3] rc1_relat_1
a: [3,196,0,2,1,118]
s: 196
i: [1,120,0,2,1,0,0,255] rc2_relat_1
a: [3,201,0,2,1,120]
s: 201
i: [1,87,0,11,1,1,0,0] t52_mcart_1
a: [3,161,0,11,1,87]
s: 161
i: [1,95,0,3,1,0,0,255] d2_xboole_0
a: [3,176,0,3,1,95]
s: 176
r: 166
i: [1,96,0,4,2,0,0,3] cc1_relat_1
s: 96
i: [1,110,0,4,1,0,0,255] existence_m1_subset_1
s: 110
d: [2,204,1,2,1,1,40,96,186]
a: [3,209,1,2,1,204]
s: 209
i: [1,111,0,4,1,0,0,2] fc6_relat_1
s: 111
d: [2,205,1,2,1,1,40,96,171]
a: [3,214,1,2,1,205]
s: 214
i: [1,97,0,5,2,0,0,1] t6_boole
s: 97
i: [1,103,0,5,2,0,0,255] t7_boole
s: 103
i: [1,100,0,6,2,0,0,255] t1_subset
s: 100
i: [1,112,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 112
i: [1,102,0,7,3,0,0,1] t8_boole
s: 102
i: [1,101,0,8,3,0,0,255] t2_subset
s: 101
i: [1,113,0,8,3,0,0,2] fc10_subset_1
s: 113
d: [2,219,1,6,2,0,40,101,110]
s: 219
i: [1,114,0,11,4,0,0,2] fc11_subset_1
s: 114
d: [2,220,2,6,2,0,40,219,112]
s: 220
i: [1,104,0,12,1,0,0,1] d4_zfmisc_1
s: 104
i: [1,115,0,14,5,0,0,1] fc12_subset_1
s: 115
d: [2,225,1,14,5,1,44,115,161]
a: [3,245,1,6,1,225]
s: 245
i: [1,105,0,16,4,0,0,2] t110_zfmisc_1
s: 105
d: [2,224,1,6,1,0,44,111,104]
s: 224
i: [1,106,0,16,4,0,0,2] t110_zfmisc_1
s: 106
i: [1,107,0,21,5,0,0,2] t32_mcart_1
s: 107
d: [2,216,1,7,3,0,44,97,97]
s: 216
i: [1,108,0,21,5,0,0,2] t32_mcart_1
s: 108
d: [2,223,1,13,3,0,44,113,104]
s: 223
i: [1,109,0,21,5,0,0,2] t32_mcart_1
s: 109
d: [2,252,1,18,4,0,44,106,104]
s: 252
d: [2,258,2,18,4,1,44,252,161]
s: 258
d: [2,248,1,21,4,0,44,105,104]
s: 248
d: [2,261,2,21,4,1,44,248,161]
s: 261
d: [2,251,1,21,4,0,44,106,104]
s: 251
d: [2,247,1,24,4,0,44,105,104]
s: 247
d: [2,264,2,21,4,1,44,251,161]
a: [3,269,2,6,1,264]
s: 269
r: 261
r: 269
a: [3,273,2,3,1,264]
s: 273
r: 258
r: 161
r: 273
a: [3,272,2,3,1,264]
s: 272
s: 161
a: [3,268,2,6,1,264]
s: 268
a: [3,276,2,12,2,264]
s: 276
d: [2,278,2,24,4,1,44,247,161]
a: [3,282,2,15,2,278]
s: 282
d: [2,339,3,14,2,1,44,276,104]
s: 339
a: [3,135,0,3,1,92]
s: 135
r: 339
r: 276
r: 161
r: 135
r: 272
a: [3,136,0,3,1,92]
s: 136
s: 273
r: 136
s: 161
d: [2,321,0,11,1,2,23,161,273]
a: [3,326,0,11,1,321]
s: 326
d: [2,358,1,13,2,1,44,326,97]
s: 358
d: [2,365,2,18,4,1,44,252,326]
d: [2,366,2,15,3,1,18,365]
s: 366
s: 258
d: [2,334,2,18,4,1,44,252,161]
s: 334
d: [2,373,2,20,5,1,44,252,358]
s: 373
s: 261
d: [2,332,2,21,4,1,44,248,161]
s: 332
d: [2,340,3,20,2,1,44,282,104]
s: 340
r: 340
r: 282
d: [2,259,2,23,4,0,44,252,104]
s: 259
d: [2,374,3,20,3,1,44,366,104]
s: 374
r: 268
r: 273
r: 374
r: 366
r: 358
r: 326
s: 136
s: 161
a: [3,229,1,2,1,225]
s: 229
r: 229
a: [3,325,0,11,1,321]
s: 325
a: [3,228,1,2,1,225]
s: 228
a: [3,233,1,2,1,225]
s: 233
r: 233
a: [3,232,1,2,1,225]
s: 232
a: [3,237,1,2,1,225]
s: 237
r: 237
r: 325
s: 326
a: [3,236,1,2,1,225]
s: 236
d: [2,436,1,13,2,1,44,326,97]
s: 436
r: 228
r: 436
r: 326
s: 325
s: 229
d: [2,464,1,13,2,1,44,325,97]
s: 464
a: [3,240,1,2,1,225]
s: 240
d: [2,416,2,2,1,1,40,229,96]
a: [3,420,2,2,1,416]
s: 420
s: 272
d: [2,286,3,18,5,1,44,107,269]
a: [3,297,3,3,1,286]
s: 297
r: 420
r: 464
r: 229
r: 325
r: 161
r: 124
r: 297
s: 124
s: 229
s: 161
s: 325
s: 420
s: 464
a: [3,296,3,3,1,286]
s: 296
d: [2,508,4,0,0,2,52,229,296,97]
e: 508
d: [2,509,4,0,0,3,52,171,229,296,102]
e: 509
d: [2,510,4,0,0,3,52,229,171,296,102]
e: 510
r: 229
r: 240
r: 272
r: 464
r: 325
s: 326
d: [2,518,1,13,2,1,44,326,97]
s: 518
s: 273
r: 161
r: 136
s: 268
d: [2,391,4,9,2,1,47,374]
d: [2,393,4,6,1,2,21,391,141]
a: [3,397,4,6,1,393]
s: 397
r: 273
r: 518
r: 326
r: 397
s: 135
s: 161
s: 240
s: 325
d: [2,576,1,2,1,2,22,240,135]
a: [3,580,1,2,1,576]
s: 580
d: [2,346,0,11,1,2,23,161,135]
a: [3,351,0,11,1,346]
s: 351
d: [2,582,1,13,2,1,44,325,97]
s: 582
s: 272
s: 276
a: [3,396,4,6,1,393]
s: 396
d: [2,588,2,21,4,1,44,251,351]
d: [2,596,2,18,3,2,21,588,141]
d: [2,597,2,12,2,3,21,596,268]
s: 597
d: [2,603,2,12,2,2,22,276,135]
s: 603
a: [3,308,3,9,2,286]
s: 308
s: 282
d: [2,288,3,18,5,1,44,108,269]
a: [3,312,3,9,2,288]
s: 312
s: 258
s: 334
d: [2,290,3,18,5,1,44,109,269]
a: [3,316,3,9,2,290]
s: 316
d: [2,407,2,18,4,1,44,252,161]
s: 407
d: [2,377,4,9,1,1,47,340]
a: [3,381,4,9,1,377]
s: 381
r: 282
d: [2,610,5,0,0,5,52,124,156,151,146,381,107]
e: 610
r: 124
r: 232
a: [3,128,0,3,1,92]
s: 128
d: [2,648,5,0,0,5,52,146,156,151,381,128,108]
e: 648
r: 128
a: [3,132,0,3,1,92]
s: 132
d: [2,650,5,0,0,5,52,146,156,151,381,132,109]
e: 650
r: 132
r: 135
r: 597
r: 268
r: 272
r: 582
r: 325
r: 396
s: 233
s: 136
d: [2,423,2,2,1,1,40,233,96]
a: [3,427,2,2,1,423]
s: 427
s: 161
s: 326
d: [2,550,5,8,4,1,44,114,397]
d: [2,560,5,6,3,2,21,550,171]
a: [3,564,5,2,1,560]
s: 564
d: [2,702,6,0,0,3,52,171,156,564,102]
e: 702
d: [2,703,6,0,0,3,52,171,156,564,102]
e: 703
d: [2,704,6,0,0,2,52,156,564,97]
e: 704
r: 564
d: [2,689,1,13,2,1,44,326,97]
s: 689
a: [3,563,5,2,1,560]
s: 563
s: 273
r: 161
r: 136
d: [2,609,2,15,2,2,23,282,381]
s: 609
a: [3,572,5,2,1,560]
s: 572
d: [2,722,6,0,0,3,52,171,146,572,102]
e: 722
d: [2,723,6,0,0,3,52,171,146,572,102]
e: 723
d: [2,724,6,0,0,2,52,146,572,97]
e: 724
r: 572
d: [2,586,2,24,4,1,44,247,351]
d: [2,592,2,21,3,2,21,586,141]
d: [2,664,2,21,3,3,22,592,381]
d: [2,665,2,21,3,3,22,664,381]
a: [3,668,2,15,2,665]
s: 668
a: [3,568,5,2,1,560]
s: 568
d: [2,730,6,0,0,3,52,171,151,568,102]
e: 730
d: [2,731,6,0,0,3,52,171,151,568,102]
e: 731
d: [2,732,6,0,0,2,52,151,568,97]
e: 732
r: 273
r: 689
r: 326
r: 381
r: 568
s: 128
s: 135
s: 240
s: 161
a: [3,567,5,2,1,560]
s: 567
s: 325
d: [2,746,1,13,2,1,44,325,97]
s: 746
a: [3,571,5,2,1,560]
s: 571
s: 272
a: [3,301,3,3,1,286]
s: 301
r: 128
r: 233
r: 312
r: 325
r: 351
r: 427
r: 746
d: [2,753,0,11,1,2,23,325,301]
a: [3,760,0,11,1,753]
s: 760
d: [2,754,0,11,1,2,23,351,301]
a: [3,765,0,11,1,754]
s: 765
s: 269
d: [2,756,1,13,2,2,23,746,301]
s: 756
d: [2,772,2,24,4,1,44,247,765]
d: [2,778,2,21,3,2,21,772,141]
a: [3,782,2,6,1,778]
s: 782
d: [2,767,1,13,2,1,44,760,97]
s: 767
d: [2,768,1,13,2,1,44,765,97]
s: 768
s: 396
d: [2,815,2,23,5,1,44,251,768]
d: [2,823,2,20,4,2,21,815,141]
a: [3,826,2,8,2,823]
s: 826
d: [2,796,3,8,2,1,44,782,97]
s: 796
d: [2,793,2,15,3,2,44,756,97]
s: 793
s: 258
d: [2,752,3,9,2,2,23,312,301]
s: 752
s: 334
d: [2,797,3,9,1,1,44,104,782]
s: 797
r: 760
r: 765
r: 756
d: [2,866,0,9,1,2,23,760,797]
a: [3,886,0,9,1,866]
s: 886
d: [2,867,0,9,1,2,23,765,797]
a: [3,891,0,9,1,867]
s: 891
r: 245
r: 767
r: 768
r: 793
d: [2,905,5,0,0,3,52,141,396,396,891,247]
e: 905
d: [2,906,5,0,0,3,52,396,891,668]
e: 906
r: 752
r: 301
r: 797
r: 796
r: 782
r: 891
s: 233
s: 128
s: 427
a: [3,890,0,9,1,867]
s: 890
s: 325
a: [3,300,3,3,1,286]
s: 300
d: [2,934,4,0,0,2,52,233,300,97]
e: 934
d: [2,935,4,0,0,3,52,171,233,300,102]
e: 935
d: [2,936,4,0,0,3,52,233,171,300,102]
e: 936
r: 233
r: 296
s: 351
s: 229
d: [2,473,0,11,1,2,23,325,297]
a: [3,479,0,11,1,473]
s: 479
s: 760
s: 245
s: 765
s: 269
d: [2,474,0,11,1,2,23,161,297]
a: [3,484,0,11,1,474]
d: [2,940,0,11,1,2,22,484,135]
a: [3,944,0,11,1,940]
s: 944
r: 245
r: 325
r: 351
r: 760
r: 765
d: [2,1002,0,11,1,2,23,325,944]
a: [3,1027,0,11,1,1002]
s: 1027
d: [2,1001,1,6,1,2,23,245,944]
a: [3,1022,1,6,1,1001]
s: 1022
d: [2,1003,0,11,1,2,23,351,944]
a: [3,1032,0,11,1,1003]
s: 1032
a: [3,781,2,6,1,778]
s: 781
d: [2,1085,3,0,0,3,52,171,781,826]
e: 1085
d: [2,1087,3,5,2,2,44,781,826]
d: [2,1088,3,2,1,2,19,1087]
d: [2,1090,3,0,0,3,21,1088,171]
e: 1090
r: 826
d: [2,1004,0,11,1,2,23,760,944]
a: [3,1037,0,11,1,1004]
s: 1037
d: [2,1005,0,11,1,2,23,765,944]
a: [3,1042,0,11,1,1005]
s: 1042
d: [2,1073,2,8,2,1,44,1022,97]
s: 1073
d: [2,1007,1,9,1,2,44,890,944]
a: [3,1052,1,9,1,1007]
s: 1052
d: [2,1086,3,8,2,1,44,781,97]
s: 1086
d: [2,1008,1,9,1,2,44,886,944]
a: [3,1057,1,9,1,1008]
s: 1057
d: [2,1006,1,11,1,2,44,479,944]
a: [3,1047,1,11,1,1006]
s: 1047
s: 312
s: 746
d: [2,989,3,9,1,1,44,104,269]
s: 989
r: 1027
r: 1032
d: [2,1117,0,9,1,2,23,1032,989]
d: [2,1128,0,0,0,3,21,1117,1052]
e: 1128
r: 229
r: 236
r: 944
r: 1042
d: [2,1146,0,9,1,2,22,351,989]
d: [2,1147,0,0,0,3,21,1146,890]
e: 1147
r: 135
s: 136
d: [2,1151,0,9,1,2,22,161,989]
a: [3,1155,0,9,1,1151]
s: 1155
r: 479
d: [2,1159,5,0,0,3,52,141,396,396,1155,247]
e: 1159
d: [2,1160,5,0,0,4,52,136,141,396,1155,251]
e: 1160
d: [2,1161,5,0,0,3,52,396,1155,668]
e: 1161
d: [2,342,4,3,1,1,47,339]
d: [2,356,0,11,1,2,22,161,273]
d: [2,677,2,21,4,1,44,251,161]
d: [2,675,2,24,4,1,44,247,161]
f: 1159,1151,675,677,732,724,704,650,648,610,560,393,377,356,342,94,87,88,89,90,91,92
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t52_mcart_1
% SZS output start Proof for mcart_1__t52_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7) => ((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t52_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7) => ((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [negated conjecture 1]
10. ! [X0] : ! [X1] : ! [X2] : ! [X3] : k4_zfmisc_1(X0,X1,X2,X3) = k2_zfmisc_1(k3_zfmisc_1(X0,X1,X2),X3) [input d4_zfmisc_1]
19. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k3_zfmisc_1(X0,X1,X2))) [input fc11_subset_1]
21. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
29. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X2,X3) => ((X1 = X3 & X0 = X2) | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t110_zfmisc_1]
32. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5) => ((X2 = X5 & X1 = X4 & X0 = X3) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t32_mcart_1]
33. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
36. ~! [X0,X1,X2,X3,X4,X5,X6,X7] : (k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7) => ((X3 = X7 & X2 = X6 & X1 = X5 & X0 = X4) | k1_xboole_0 = X3 | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 2]
41. ! [X0,X1,X2,X3] : k4_zfmisc_1(X0,X1,X2,X3) = k2_zfmisc_1(k3_zfmisc_1(X0,X1,X2),X3) [flattening 10]
42. ! [X0,X1,X2,X3] : (k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X2,X3) => ((X1 = X3 & X0 = X2) | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 29]
43. ! [X0,X1,X2,X3,X4,X5] : (k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(X3,X4,X5) => ((X2 = X5 & X1 = X4 & X0 = X3) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 32]
51. ? [X0,X1,X2,X3,X4,X5,X6,X7] : (((X3 != X7 | X2 != X6 | X1 != X5 | X0 != X4) & k1_xboole_0 != X3 & k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) & k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7)) [ennf transformation 36]
52. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((X3 != X7 | X2 != X6 | X1 != X5 | X0 != X4) & k1_xboole_0 != X3 & k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0 & k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7)) [flattening 51]
54. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 33]
64. ! [X0,X1,X2,X3] : (((X1 = X3 & X0 = X2) | k1_xboole_0 = X1 | k1_xboole_0 = X0) | k2_zfmisc_1(X0,X1) != k2_zfmisc_1(X2,X3)) [ennf transformation 42]
65. ! [X0,X1,X2,X3] : ((X1 = X3 & X0 = X2) | k1_xboole_0 = X1 | k1_xboole_0 = X0 | k2_zfmisc_1(X0,X1) != k2_zfmisc_1(X2,X3)) [flattening 64]
66. ! [X0,X1,X2,X3,X4,X5] : (((X2 = X5 & X1 = X4 & X0 = X3) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0) | k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5)) [ennf transformation 43]
67. ! [X0,X1,X2,X3,X4,X5] : ((X2 = X5 & X1 = X4 & X0 = X3) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5)) [flattening 66]
71. ! [X0,X1,X2] : (~v1_xboole_0(k3_zfmisc_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 19]
72. ! [X0,X1,X2] : (~v1_xboole_0(k3_zfmisc_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 71]
75. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((X3 != X7 | X2 != X6 | X1 != X5 | X0 != X4) & k1_xboole_0 != X3 & k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0 & k4_zfmisc_1(X0,X1,X2,X3) = k4_zfmisc_1(X4,X5,X6,X7)) => ((sK3 != sK7 | sK2 != sK6 | sK1 != sK5 | sK0 != sK4) & k1_xboole_0 != sK3 & k1_xboole_0 != sK2 & k1_xboole_0 != sK1 & k1_xboole_0 != sK0 & k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7)) [choice axiom]
76. (sK3 != sK7 | sK2 != sK6 | sK1 != sK5 | sK0 != sK4) & k1_xboole_0 != sK3 & k1_xboole_0 != sK2 & k1_xboole_0 != sK1 & k1_xboole_0 != sK0 & k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7) [skolemisation 52,75]
87. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7) [cnf transformation 76] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
88. k1_xboole_0 != sK0 [cnf transformation 76] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
89. k1_xboole_0 != sK1 [cnf transformation 76] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
90. k1_xboole_0 != sK2 [cnf transformation 76] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
91. k1_xboole_0 != sK3 [cnf transformation 76] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
92. sK3 != sK7 | sK2 != sK6 | sK1 != sK5 | sK0 != sK4 [cnf transformation 76] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
94. v1_xboole_0(k1_xboole_0) [cnf transformation 21] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
97. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 54] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. k4_zfmisc_1(X0,X1,X2,X3) = k2_zfmisc_1(k3_zfmisc_1(X0,X1,X2),X3) [cnf transformation 41] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
105. k2_zfmisc_1(X0,X1) != k2_zfmisc_1(X2,X3) | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X0 = X2 [cnf transformation 65] {a:0,w:16,wCS:6400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
106. k2_zfmisc_1(X0,X1) != k2_zfmisc_1(X2,X3) | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X1 = X3 [cnf transformation 65] {a:0,w:16,wCS:6400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
107. k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X0 = X3 [cnf transformation 67] {a:0,w:21,wCS:8400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
108. k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X1 = X4 [cnf transformation 67] {a:0,w:21,wCS:8400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
109. k3_zfmisc_1(X0,X1,X2) != k3_zfmisc_1(X3,X4,X5) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0 | X2 = X5 [cnf transformation 67] {a:0,w:21,wCS:8400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
114. ~v1_xboole_0(k3_zfmisc_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 72] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
122. 1 <=> sK0 = sK4 [avatar definition]
124. sK0 != sK4 <- (~1) [avatar component clause 122] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
126. 2 <=> sK1 = sK5 [avatar definition]
128. sK1 != sK5 <- (~2) [avatar component clause 126] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
130. 3 <=> sK2 = sK6 [avatar definition]
132. sK2 != sK6 <- (~3) [avatar component clause 130] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
134. 4 <=> sK3 = sK7 [avatar definition]
137. ~1 | ~2 | ~3 | ~4 [avatar split clause 92,134,130,126,122]
139. 5 <=> k1_xboole_0 = sK3 [avatar definition]
141. k1_xboole_0 != sK3 <- (~5) [avatar component clause 139] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
142. ~5 [avatar split clause 91,139]
144. 6 <=> k1_xboole_0 = sK2 [avatar definition]
146. k1_xboole_0 != sK2 <- (~6) [avatar component clause 144] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
147. ~6 [avatar split clause 90,144]
149. 7 <=> k1_xboole_0 = sK1 [avatar definition]
151. k1_xboole_0 != sK1 <- (~7) [avatar component clause 149] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
152. ~7 [avatar split clause 89,149]
154. 8 <=> k1_xboole_0 = sK0 [avatar definition]
156. k1_xboole_0 != sK0 <- (~8) [avatar component clause 154] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
157. ~8 [avatar split clause 88,154]
159. 9 <=> k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7) [avatar definition]
161. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,sK7) <- (9) [avatar component clause 159] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
162. 9 [avatar split clause 87,159]
169. 11 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
171. v1_xboole_0(k1_xboole_0) <- (11) [avatar component clause 169] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
172. 11 [avatar split clause 94,169]
247. k4_zfmisc_1(X0,X1,X2,X3) != k2_zfmisc_1(X4,X5) | k1_xboole_0 = X3 | k1_xboole_0 = k3_zfmisc_1(X0,X1,X2) | k3_zfmisc_1(X0,X1,X2) = X4 [superposition 105,104] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
251. k4_zfmisc_1(X0,X1,X2,X3) != k2_zfmisc_1(X4,X5) | k1_xboole_0 = X3 | k1_xboole_0 = k3_zfmisc_1(X0,X1,X2) | X3 = X5 [superposition 106,104] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
252. k4_zfmisc_1(X0,X1,X2,X3) != k2_zfmisc_1(X4,X5) | k1_xboole_0 = X5 | k1_xboole_0 = X4 | X3 = X5 [superposition 106,104] {a:1,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
267. 25 <=> k1_xboole_0 = k3_zfmisc_1(sK4,sK5,sK6) [avatar definition]
269. k1_xboole_0 = k3_zfmisc_1(sK4,sK5,sK6) <- (25) [avatar component clause 267] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
271. 26 <=> k1_xboole_0 = sK7 [avatar definition]
273. k1_xboole_0 = sK7 <- (26) [avatar component clause 271] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
275. 27 <=> ! [X1,X0] : (k2_zfmisc_1(X0,X1) != k4_zfmisc_1(sK0,sK1,sK2,sK3) | sK7 = X1) [avatar definition]
276. k2_zfmisc_1(X0,X1) != k4_zfmisc_1(sK0,sK1,sK2,sK3) | sK7 = X1 <- (27) [avatar component clause 275] {a:2,w:12,wCS:1200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
281. 28 <=> ! [X1,X0] : (k2_zfmisc_1(X0,X1) != k4_zfmisc_1(sK0,sK1,sK2,sK3) | k3_zfmisc_1(sK4,sK5,sK6) = X0) [avatar definition]
282. k2_zfmisc_1(X0,X1) != k4_zfmisc_1(sK0,sK1,sK2,sK3) | k3_zfmisc_1(sK4,sK5,sK6) = X0 <- (28) [avatar component clause 281] {a:2,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
324. 35 <=> k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,k1_xboole_0) [avatar definition]
326. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,k1_xboole_0) <- (35) [avatar component clause 324] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
339. k4_zfmisc_1(X0,X1,X2,X3) != k4_zfmisc_1(sK0,sK1,sK2,sK3) | sK7 = X3 <- (27) [superposition 276,104] {a:3,w:14,wCS:1400,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
340. k4_zfmisc_1(X0,X1,X2,X3) != k4_zfmisc_1(sK0,sK1,sK2,sK3) | k3_zfmisc_1(X0,X1,X2) = k3_zfmisc_1(sK4,sK5,sK6) <- (28) [superposition 282,104] {a:3,w:20,wCS:2000,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
342. sK3 = sK7 <- (27) [equality resolution 339] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
343. 4 | ~27 [avatar split clause 342,275,134]
356. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k4_zfmisc_1(sK4,sK5,sK6,k1_xboole_0) <- (9, 26) [forward demodulation 161,273] {a:0,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
357. 35 | ~9 | ~26 [avatar split clause 356,271,159,324]
365. k4_zfmisc_1(sK0,sK1,sK2,sK3) != k2_zfmisc_1(X6,X7) | k1_xboole_0 = X7 | k1_xboole_0 = X6 | k1_xboole_0 = X7 <- (35) [superposition 252,326] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
366. k4_zfmisc_1(sK0,sK1,sK2,sK3) != k2_zfmisc_1(X6,X7) | k1_xboole_0 = X7 | k1_xboole_0 = X6 <- (35) [duplicate literal removal 365] {a:2,w:15,wCS:1500,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
374. k4_zfmisc_1(X0,X1,X2,X3) != k4_zfmisc_1(sK0,sK1,sK2,sK3) | k1_xboole_0 = X3 | k1_xboole_0 = k3_zfmisc_1(X0,X1,X2) <- (35) [superposition 366,104] {a:3,w:20,wCS:2000,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
377. k3_zfmisc_1(sK4,sK5,sK6) = k3_zfmisc_1(sK0,sK1,sK2) <- (28) [equality resolution 340] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
379. 37 <=> k3_zfmisc_1(sK4,sK5,sK6) = k3_zfmisc_1(sK0,sK1,sK2) [avatar definition]
381. k3_zfmisc_1(sK4,sK5,sK6) = k3_zfmisc_1(sK0,sK1,sK2) <- (37) [avatar component clause 379] {a:4,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
382. 37 | ~28 [avatar split clause 377,281,379]
391. k1_xboole_0 = sK3 | k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) <- (35) [equality resolution 374] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
393. k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) <- (~5, 35) [subsumption resolution 391,141] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
395. 38 <=> k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) [avatar definition]
396. k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) <- (~38) [avatar component clause 395] {a:4,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
397. k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) <- (38) [avatar component clause 395] {a:4,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
398. 38 | 5 | ~35 [avatar split clause 393,324,139,395]
550. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (38) [superposition 114,397] {a:5,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
560. v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (11, 38) [subsumption resolution 550,171] {a:5,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
562. 47 <=> v1_xboole_0(sK0) [avatar definition]
564. v1_xboole_0(sK0) <- (47) [avatar component clause 562] {a:5,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
566. 48 <=> v1_xboole_0(sK1) [avatar definition]
568. v1_xboole_0(sK1) <- (48) [avatar component clause 566] {a:5,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
570. 49 <=> v1_xboole_0(sK2) [avatar definition]
572. v1_xboole_0(sK2) <- (49) [avatar component clause 570] {a:5,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
573. 47 | 48 | 49 | ~11 | ~38 [avatar split clause 560,395,169,570,566,562]
610. $false <- (~1, ~6, ~7, ~8, 37) [unit resulting resolution 124,156,151,146,381,107] {a:5,w:0,goal:1,sine:2,thAx:0,allAx:14,thDist:-14}
623. 1 | 6 | 7 | 8 | ~37 [avatar contradiction clause 610]
648. $false <- (~2, ~6, ~7, ~8, 37) [unit resulting resolution 146,156,151,381,128,108] {a:5,w:0,goal:1,sine:2,thAx:0,allAx:14,thDist:-14}
649. 2 | 6 | 7 | 8 | ~37 [avatar contradiction clause 648]
650. $false <- (~3, ~6, ~7, ~8, 37) [unit resulting resolution 146,156,151,381,132,109] {a:5,w:0,goal:1,sine:2,thAx:0,allAx:14,thDist:-14}
651. 3 | 6 | 7 | 8 | ~37 [avatar contradiction clause 650]
675. k2_zfmisc_1(X0,X1) != k4_zfmisc_1(sK0,sK1,sK2,sK3) | k1_xboole_0 = sK7 | k1_xboole_0 = k3_zfmisc_1(sK4,sK5,sK6) | k3_zfmisc_1(sK4,sK5,sK6) = X0 <- (9) [superposition 247,161] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
677. k4_zfmisc_1(sK0,sK1,sK2,sK3) != k2_zfmisc_1(X4,X5) | k1_xboole_0 = sK7 | k1_xboole_0 = k3_zfmisc_1(sK4,sK5,sK6) | sK7 = X5 <- (9) [superposition 251,161] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
704. $false <- (~8, 47) [unit resulting resolution 156,564,97] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
708. 8 | ~47 [avatar contradiction clause 704]
724. $false <- (~6, 49) [unit resulting resolution 146,572,97] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
728. 6 | ~49 [avatar contradiction clause 724]
732. $false <- (~7, 48) [unit resulting resolution 151,568,97] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
736. 7 | ~48 [avatar contradiction clause 732]
738. 25 | 26 | 27 | ~9 [avatar split clause 677,159,275,271,267]
740. 25 | 26 | 28 | ~9 [avatar split clause 675,159,281,271,267]
989. k4_zfmisc_1(sK4,sK5,sK6,X0) = k2_zfmisc_1(k1_xboole_0,X0) <- (25) [superposition 104,269] {a:3,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1151. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k2_zfmisc_1(k1_xboole_0,sK7) <- (9, 25) [forward demodulation 161,989] {a:0,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1153. 77 <=> k4_zfmisc_1(sK0,sK1,sK2,sK3) = k2_zfmisc_1(k1_xboole_0,sK7) [avatar definition]
1155. k4_zfmisc_1(sK0,sK1,sK2,sK3) = k2_zfmisc_1(k1_xboole_0,sK7) <- (77) [avatar component clause 1153] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1156. 77 | ~9 | ~25 [avatar split clause 1151,267,159,1153]
1159. $false <- (~5, ~38, 77) [unit resulting resolution 141,396,396,1155,247] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:16,thDist:-16}
1174. 5 | 38 | ~77 [avatar contradiction clause 1159]
1188. $false [avatar sat refutation 137,142,147,152,157,162,172,343,357,382,398,573,623,649,651,708,728,736,738,740,1156,1174] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1156,thDist:-1156}
% SZS output end Proof for mcart_1__t52_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 34
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 34
% Generated clauses: 834
% Active clauses: 170
% Passive clauses: 312
% Final active clauses: 68
% Final passive clauses: 50
% Main loop iterations started: 191
% 
% >>> Simplifying Inferences
% Duplicate literals: 17
% Trivial inequalities: 12
% Fw subsumption resolutions: 93
% Fw demodulations: 35
% Bw demodulations: 41
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Equational tautologies: 6
% Forward subsumptions: 175
% 
% >>> Generating Inferences
% Binary resolution: 17
% Unit resulting resolution: 24
% Forward superposition: 68
% Backward superposition: 274
% Self superposition: 1
% Equality resolution: 9
% 
% >>> AVATAR
% Split clauses: 61
% Split components: 162
% Sat splitting refutations: 27
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 155
% SAT solver unit clauses: 13
% SAT solver binary clauses: 21
% 
% Memory used [KB]: 5756
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% sine selection: 0.001 s
% splitting model update: 0.001 s ( own 0.0000 s ) 
% backward superposition index maintenance: 0.001 s
% forward superposition index maintenance: 0.001 s
% forward subsumption: 0.002 s
% unit resulting resolution: 0.006 s
% superposition: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.017 s ( own 0.003 s ) 

