i: [1,68,0,3,1,1,0,0] t41_subset_1
a: [3,113,0,3,1,68]
s: 113
i: [1,69,0,5,1,1,0,0] t41_subset_1
a: [3,108,0,5,1,69]
s: 108
i: [1,70,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,118,0,2,1,70]
s: 118
i: [1,71,0,2,1,0,0,255] fc1_xboole_0
a: [3,123,0,2,1,71]
s: 123
i: [1,98,0,2,1,0,0,2] rc2_zfmisc_1
a: [3,133,0,2,1,98]
s: 133
i: [1,99,0,2,1,0,0,2] rc2_xboole_0
a: [3,138,0,2,1,99]
s: 138
i: [1,100,0,2,1,0,0,2] rc1_xboole_0
a: [3,143,0,2,1,100]
s: 143
i: [1,102,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,148,0,2,1,102]
s: 148
i: [1,101,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,153,0,2,1,101]
s: 153
i: [1,73,0,3,1,0,0,1] fc2_xboole_0
s: 73
i: [1,74,0,3,1,0,0,1] fc1_subset_1
s: 74
i: [1,75,0,3,1,0,0,1] fc2_zfmisc_1
s: 75
i: [1,95,0,3,1,0,0,1] rc2_subset_1
s: 95
i: [1,96,0,3,1,0,0,255] reflexivity_r1_tarski
s: 96
i: [1,72,0,3,1,0,0,255] d2_xboole_0
a: [3,128,0,3,1,72]
s: 128
r: 118
i: [1,76,0,4,2,0,0,2] cc2_zfmisc_1
s: 76
d: [2,156,1,2,1,1,40,76,123]
a: [3,162,1,2,1,156]
s: 162
i: [1,79,0,4,2,0,0,2] cc1_zfmisc_1
s: 79
d: [2,158,1,2,1,1,40,76,143]
a: [3,167,1,2,1,158]
s: 167
i: [1,93,0,4,1,0,0,1] existence_m1_subset_1
s: 93
i: [1,78,0,5,2,0,0,1] rc1_subset_1
s: 78
d: [2,157,1,3,1,0,40,76,95]
s: 157
i: [1,80,0,5,2,0,0,255] t6_boole
s: 80
i: [1,92,0,5,2,0,0,1] t7_boole
s: 92
i: [1,94,0,5,1,0,0,1] rc2_subset_1
s: 94
i: [1,97,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 97
d: [2,172,1,2,1,1,40,92,113]
a: [3,176,1,2,1,172]
s: 176
i: [1,77,0,7,2,0,0,1] rc1_subset_1
s: 77
d: [2,178,1,3,1,1,40,97,113]
a: [3,182,1,3,1,178]
s: 182
i: [1,83,0,7,3,0,0,1] d1_subset_1
s: 83
i: [1,84,0,7,3,0,0,1] d1_subset_1
s: 84
d: [2,184,1,5,2,0,40,83,93]
s: 184
i: [1,89,0,7,2,0,0,255] t31_zfmisc_1
s: 89
d: [2,191,1,4,1,0,40,89,96]
s: 191
i: [1,90,0,7,2,0,0,255] t31_zfmisc_1
s: 90
i: [1,91,0,7,3,0,0,255] t8_boole
s: 91
d: [2,192,2,4,1,0,40,191,97]
s: 192
i: [1,103,0,7,2,0,0,255] d1_zfmisc_1
s: 103
d: [2,190,2,5,2,0,40,184,76]
s: 190
i: [1,104,0,7,2,0,0,255] d1_zfmisc_1
s: 104
i: [1,81,0,8,3,0,0,1] d1_subset_1
s: 81
d: [2,202,1,8,2,0,40,81,94]
d: [2,204,1,5,1,0,21,202,74]
s: 204
i: [1,82,0,8,3,0,0,1] d1_subset_1
s: 82
d: [2,205,2,4,1,0,40,204,104]
s: 205
d: [2,208,0,6,2,0,21,82,92]
s: 208
i: [1,87,0,14,3,0,0,255] d1_zfmisc_1
s: 87
d: [2,209,1,3,1,1,40,208,113]
a: [3,216,1,3,1,209]
s: 216
i: [1,88,0,14,3,0,0,255] d1_zfmisc_1
s: 88
d: [2,210,2,4,1,0,40,208,191]
s: 210
d: [2,201,1,6,2,0,40,81,93]
s: 201
d: [2,169,1,7,3,0,44,80,80]
s: 169
d: [2,206,2,5,1,0,40,204,97]
s: 206
d: [2,195,1,7,2,0,40,103,97]
s: 195
d: [2,227,2,8,2,0,40,201,104]
d: [2,228,2,5,1,0,21,227,74]
s: 228
d: [2,200,1,10,3,0,40,81,77]
d: [2,203,1,7,2,0,21,200,74]
s: 203
d: [2,211,1,7,2,0,40,208,103]
s: 211
d: [2,235,2,4,1,1,40,211,108]
a: [3,241,2,4,1,235]
s: 241
d: [2,243,3,0,0,2,52,113,241,90]
e: 243
d: [2,244,3,3,1,1,40,241,90]
d: [2,246,3,0,0,2,21,244,113]
e: 246
f: 246,235,68,69
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t41_subset_1
% SZS output start Proof for subset_1__t41_subset_1
1. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(k1_tarski(X0),k1_zfmisc_1(X1))) [input t41_subset_1]
2. ~! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(k1_tarski(X0),k1_zfmisc_1(X1))) [negated conjecture 1]
6. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
7. ! [X0] : ! [X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [input d1_zfmisc_1]
26. ! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [input t31_zfmisc_1]
28. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
30. ~! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(k1_tarski(X0),k1_zfmisc_1(X1))) [flattening 2]
31. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 6]
32. ! [X0,X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [flattening 7]
33. ! [X0,X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [flattening 26]
35. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 28]
37. ? [X0,X1] : (~m1_subset_1(k1_tarski(X0),k1_zfmisc_1(X1)) & r2_hidden(X0,X1)) [ennf transformation 30]
42. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 31]
44. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 35]
46. ? [X0,X1] : (~m1_subset_1(k1_tarski(X0),k1_zfmisc_1(X1)) & r2_hidden(X0,X1)) => (~m1_subset_1(k1_tarski(sK0),k1_zfmisc_1(sK1)) & r2_hidden(sK0,sK1)) [choice axiom]
47. ~m1_subset_1(k1_tarski(sK0),k1_zfmisc_1(sK1)) & r2_hidden(sK0,sK1) [skolemisation 37,46]
50. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 42]
51. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ~r1_tarski(X2,X0)) & (r1_tarski(X2,X0) | ~r2_hidden(X2,X1))) | k1_zfmisc_1(X0) != X1)) [nnf transformation 32]
52. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [rectify 51]
53. ! [X1,X0] : (? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1))) => ((~r1_tarski(sK3(X0,X1),X0) | ~r2_hidden(sK3(X0,X1),X1)) & (r1_tarski(sK3(X0,X1),X0) | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
54. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ((~r1_tarski(sK3(X0,X1),X0) | ~r2_hidden(sK3(X0,X1),X1)) & (r1_tarski(sK3(X0,X1),X0) | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [skolemisation 52,53]
55. ! [X0,X1] : ((r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X1) | ~r1_tarski(k1_tarski(X0),X1))) [nnf transformation 33]
68. r2_hidden(sK0,sK1) [cnf transformation 47] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. ~m1_subset_1(k1_tarski(sK0),k1_zfmisc_1(sK1)) [cnf transformation 47] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
82. m1_subset_1(X1,X0) | ~r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 50] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
86. r2_hidden(X3,X1) | ~r1_tarski(X3,X0) | k1_zfmisc_1(X0) != X1 [cnf transformation 54] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
90. r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 55] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
92. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 44] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
103. r2_hidden(X3,k1_zfmisc_1(X0)) | ~r1_tarski(X3,X0) [equality resolution 86] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
106. 1 <=> m1_subset_1(k1_tarski(sK0),k1_zfmisc_1(sK1)) [avatar definition]
108. ~m1_subset_1(k1_tarski(sK0),k1_zfmisc_1(sK1)) <- (~1) [avatar component clause 106] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
109. ~1 [avatar split clause 69,106]
111. 2 <=> r2_hidden(sK0,sK1) [avatar definition]
113. r2_hidden(sK0,sK1) <- (2) [avatar component clause 111] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
114. 2 [avatar split clause 68,111]
208. ~r2_hidden(X1,X0) | m1_subset_1(X1,X0) [subsumption resolution 82,92] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
211. m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r1_tarski(X1,X2) [resolution 208,103] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
235. ~r1_tarski(k1_tarski(sK0),sK1) <- (~1) [resolution 211,108] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
239. 16 <=> r1_tarski(k1_tarski(sK0),sK1) [avatar definition]
241. ~r1_tarski(k1_tarski(sK0),sK1) <- (~16) [avatar component clause 239] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
242. ~16 | 1 [avatar split clause 235,106,239]
244. ~r2_hidden(sK0,sK1) <- (~16) [resolution 241,90] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
246. $false <- (2, ~16) [subsumption resolution 244,113] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
247. ~2 | 16 [avatar contradiction clause 246]
248. $false [avatar sat refutation 109,114,242,247] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:72,thDist:-72}
% SZS output end Proof for subset_1__t41_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 28
% Introduced skolems: 10
% 
% >>> Saturation
% Initial clauses: 35
% Generated clauses: 112
% Active clauses: 54
% Passive clauses: 63
% Final active clauses: 53
% Final passive clauses: 6
% Main loop iterations started: 57
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 20
% 
% >>> Generating Inferences
% Binary resolution: 49
% Unit resulting resolution: 1
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 10
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% unit resulting resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.001 s ( own 0.0000 s ) 

