i: [1,51,0,2,1,1,0,0] t132_zfmisc_1
a: [3,91,0,2,1,51]
s: 91
i: [1,52,0,3,1,1,0,0] t132_zfmisc_1
a: [3,86,0,3,1,52]
s: 86
i: [1,53,0,4,1,1,0,0] t132_zfmisc_1
a: [3,81,0,4,1,53]
s: 81
i: [1,54,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,96,0,2,1,54]
s: 96
i: [1,55,0,2,1,0,0,255] fc1_xboole_0
a: [3,101,0,2,1,55]
s: 101
i: [1,70,0,2,1,0,0,2] rc2_xboole_0
a: [3,111,0,2,1,70]
s: 111
i: [1,71,0,2,1,0,0,1] rc2_zfmisc_1
a: [3,116,0,2,1,71]
s: 116
i: [1,72,0,2,1,0,0,2] rc1_xboole_0
a: [3,121,0,2,1,72]
s: 121
i: [1,74,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,126,0,2,1,74]
s: 126
i: [1,73,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,131,0,2,1,73]
s: 131
i: [1,57,0,3,1,0,0,1] fc2_xboole_0
s: 57
i: [1,58,0,3,1,0,0,1] fc2_zfmisc_1
s: 58
i: [1,56,0,3,1,0,0,255] d2_xboole_0
a: [3,106,0,3,1,56]
s: 106
r: 96
i: [1,59,0,4,2,0,0,1] cc2_zfmisc_1
s: 59
i: [1,60,0,4,2,0,0,1] cc1_zfmisc_1
s: 60
i: [1,76,0,4,1,0,0,1] d1_tarski
s: 76
d: [2,134,1,2,1,1,40,59,121]
a: [3,139,1,2,1,134]
s: 139
i: [1,61,0,5,2,0,0,255] t6_boole
s: 61
d: [2,135,1,2,1,1,40,59,101]
a: [3,144,1,2,1,135]
s: 144
i: [1,68,0,5,2,0,0,1] t7_boole
s: 68
i: [1,69,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 69
d: [2,149,1,2,1,1,40,68,86]
a: [3,154,1,2,1,149]
s: 154
i: [1,67,0,7,3,0,0,1] t8_boole
s: 67
d: [2,156,1,3,1,1,40,69,86]
a: [3,161,1,3,1,156]
s: 161
i: [1,77,0,7,2,0,0,1] d1_tarski
s: 77
i: [1,62,0,9,3,0,0,1] t131_zfmisc_1
s: 62
d: [2,157,1,4,1,0,40,69,76]
s: 157
i: [1,65,0,14,3,0,0,1] d1_tarski
s: 65
d: [2,146,1,7,3,0,44,61,61]
s: 146
i: [1,66,0,14,3,0,0,1] d1_tarski
s: 66
d: [2,165,1,8,3,0,44,76,62]
s: 165
d: [2,175,2,11,3,0,40,165,77]
d: [2,176,2,8,2,0,21,175,58]
d: [2,177,2,5,1,0,21,176,57]
s: 177
d: [2,164,1,11,4,0,44,77,62]
d: [2,168,1,9,3,0,21,164,68]
s: 168
d: [2,169,2,8,3,0,44,157,62]
s: 169
d: [2,171,1,11,3,0,40,65,68]
s: 171
d: [2,170,1,14,3,0,40,65,69]
s: 170
d: [2,173,2,8,3,0,40,165,69]
s: 173
d: [2,172,1,17,3,0,40,65,77]
s: 172
d: [2,182,2,13,5,0,44,62,168]
d: [2,184,2,11,4,0,18,182]
d: [2,187,2,9,3,0,21,184,68]
s: 187
d: [2,209,3,0,0,3,52,91,86,81,187]
e: 209
f: 209,51,52,53
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t132_zfmisc_1
% SZS output start Proof for zfmisc_1__t132_zfmisc_1
1. ! [X0] : ! [X1] : (v1_zfmisc_1(X1) => (r2_hidden(X0,X1) => k1_tarski(X0) = X1)) [input t132_zfmisc_1]
2. ~! [X0] : ! [X1] : (v1_zfmisc_1(X1) => (r2_hidden(X0,X1) => k1_tarski(X0) = X1)) [negated conjecture 1]
6. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
18. ! [X0] : ~(! [X1] : k1_tarski(X1) != X0 & v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) [input t131_zfmisc_1]
20. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
22. ~! [X0,X1] : (v1_zfmisc_1(X1) => (r2_hidden(X0,X1) => k1_tarski(X0) = X1)) [flattening 2]
23. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 6]
25. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 20]
26. ? [X0,X1] : ((k1_tarski(X0) != X1 & r2_hidden(X0,X1)) & v1_zfmisc_1(X1)) [ennf transformation 22]
27. ? [X0,X1] : (k1_tarski(X0) != X1 & r2_hidden(X0,X1) & v1_zfmisc_1(X1)) [flattening 26]
31. ! [X0] : (? [X1] : k1_tarski(X1) = X0 | ~v1_zfmisc_1(X0) | v1_xboole_0(X0)) [ennf transformation 18]
33. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 25]
35. ? [X0,X1] : (k1_tarski(X0) != X1 & r2_hidden(X0,X1) & v1_zfmisc_1(X1)) => (sK1 != k1_tarski(sK0) & r2_hidden(sK0,sK1) & v1_zfmisc_1(sK1)) [choice axiom]
36. sK1 != k1_tarski(sK0) & r2_hidden(sK0,sK1) & v1_zfmisc_1(sK1) [skolemisation 27,35]
37. ! [X0] : (? [X1] : k1_tarski(X1) = X0 => k1_tarski(sK2(X0)) = X0) [choice axiom]
38. ! [X0] : (k1_tarski(sK2(X0)) = X0 | ~v1_zfmisc_1(X0) | v1_xboole_0(X0)) [skolemisation 31,37]
39. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 23]
40. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 39]
41. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 40,41]
51. v1_zfmisc_1(sK1) [cnf transformation 36] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. r2_hidden(sK0,sK1) [cnf transformation 36] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. sK1 != k1_tarski(sK0) [cnf transformation 36] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
62. k1_tarski(sK2(X0)) = X0 | ~v1_zfmisc_1(X0) | v1_xboole_0(X0) [cnf transformation 38] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
63. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 42] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
68. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 33] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
77. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 63] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
79. 1 <=> sK1 = k1_tarski(sK0) [avatar definition]
81. sK1 != k1_tarski(sK0) <- (~1) [avatar component clause 79] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
82. ~1 [avatar split clause 53,79]
84. 2 <=> r2_hidden(sK0,sK1) [avatar definition]
86. r2_hidden(sK0,sK1) <- (2) [avatar component clause 84] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. 2 [avatar split clause 52,84]
89. 3 <=> v1_zfmisc_1(sK1) [avatar definition]
91. v1_zfmisc_1(sK1) <- (3) [avatar component clause 89] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
92. 3 [avatar split clause 51,89]
164. ~r2_hidden(X1,X0) | sK2(X0) = X1 | ~v1_zfmisc_1(X0) | v1_xboole_0(X0) [superposition 77,62] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:4,thDist:-4}
168. sK2(X0) = X1 | ~r2_hidden(X1,X0) | ~v1_zfmisc_1(X0) [subsumption resolution 164,68] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
182. k1_tarski(X3) = X2 | ~v1_zfmisc_1(X2) | v1_xboole_0(X2) | ~r2_hidden(X3,X2) | ~v1_zfmisc_1(X2) [superposition 62,168] {a:2,w:13,wCS:5200,sine:1,thAx:0,allAx:6,thDist:-6}
184. k1_tarski(X3) = X2 | ~v1_zfmisc_1(X2) | v1_xboole_0(X2) | ~r2_hidden(X3,X2) [duplicate literal removal 182] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:6,thDist:-6}
187. k1_tarski(X3) = X2 | ~v1_zfmisc_1(X2) | ~r2_hidden(X3,X2) [subsumption resolution 184,68] {a:2,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
209. $false <- (~1, 2, 3) [unit resulting resolution 91,86,81,187] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
221. 1 | ~2 | ~3 [avatar contradiction clause 209]
222. $false [avatar sat refutation 82,87,92,221] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:61,thDist:-61}
% SZS output end Proof for zfmisc_1__t132_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 104
% Active clauses: 36
% Passive clauses: 43
% Final active clauses: 35
% Final passive clauses: 4
% Main loop iterations started: 39
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Trivial inequalities: 1
% Fw subsumption resolutions: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Equational tautologies: 1
% Forward subsumptions: 18
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 1
% Forward superposition: 7
% Backward superposition: 23
% Self superposition: 3
% Equality factoring: 1
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 11
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

