i: [1,65,0,2,1,1,0,0] t4_xtuple_0
a: [3,119,0,2,1,65]
s: 119
i: [1,64,0,2,1,1,0,0] t4_xtuple_0
a: [3,124,0,2,1,64]
s: 124
i: [1,69,0,3,1,1,0,0] t4_xtuple_0
a: [3,99,0,3,1,69]
s: 99
i: [1,68,0,5,1,1,0,0] t4_xtuple_0
a: [3,104,0,5,1,68]
s: 104
i: [1,67,0,5,1,1,0,0] t4_xtuple_0
a: [3,109,0,5,1,67]
s: 109
i: [1,66,0,5,1,1,0,0] t4_xtuple_0
a: [3,114,0,5,1,66]
s: 114
i: [1,70,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,129,0,2,1,70]
s: 129
i: [1,71,0,2,1,0,0,255] fc1_xboole_0
a: [3,134,0,2,1,71]
s: 134
i: [1,92,0,2,1,0,0,255] rc2_xboole_0
a: [3,144,0,2,1,92]
s: 144
i: [1,93,0,2,1,0,0,3] rc1_xtuple_0
a: [3,149,0,2,1,93]
s: 149
i: [1,94,0,2,1,0,0,255] rc1_xboole_0
a: [3,154,0,2,1,94]
s: 154
i: [1,95,0,2,1,0,0,1] rc2_xtuple_0
a: [3,159,0,2,1,95]
s: 159
i: [1,73,0,3,1,0,0,255] fc2_xboole_0
s: 73
i: [1,72,0,3,1,0,0,255] d2_xboole_0
a: [3,139,0,3,1,72]
s: 139
r: 129
i: [1,78,0,4,2,0,0,3] cc1_xtuple_0
s: 78
i: [1,83,0,4,1,0,0,255] fc3_xboole_0
s: 83
d: [2,162,1,2,1,1,40,78,119]
a: [3,168,1,2,1,162]
s: 168
i: [1,84,0,4,1,0,0,3] fc1_xtuple_0
s: 84
d: [2,163,1,2,1,1,40,78,124]
a: [3,173,1,2,1,163]
s: 173
i: [1,77,0,5,2,0,0,255] t6_boole
s: 77
i: [1,88,0,5,1,0,0,1] fc2_xtuple_0
s: 88
d: [2,164,1,2,1,1,40,78,159]
a: [3,178,1,2,1,164]
s: 178
i: [1,74,0,6,1,0,0,1] d6_xtuple_0
s: 74
d: [2,182,1,5,1,0,40,88,78]
s: 182
i: [1,75,0,6,1,0,0,1] d7_xtuple_0
s: 75
i: [1,85,0,6,1,0,0,2] rd1_xtuple_0
s: 85
i: [1,86,0,6,1,0,0,1] rd2_xtuple_0
s: 86
i: [1,81,0,7,3,0,0,255] t8_boole
s: 81
i: [1,87,0,7,1,0,0,255] commutativity_k2_tarski
s: 87
i: [1,89,0,7,1,0,0,1] rd5_xtuple_0
s: 89
i: [1,90,0,7,1,0,0,1] rd4_xtuple_0
s: 90
i: [1,91,0,7,1,0,0,1] rd6_xtuple_0
s: 91
i: [1,76,0,9,2,0,0,2] rd3_xtuple_0
s: 76
d: [2,194,1,9,2,1,44,76,104]
d: [2,201,1,7,1,2,21,194,168]
a: [3,205,1,7,1,201]
s: 205
i: [1,80,0,10,1,0,0,255] d5_tarski
s: 80
i: [1,82,0,10,1,0,0,2] d4_xtuple_0
s: 82
d: [2,218,2,10,1,1,44,82,205]
s: 218
i: [1,79,0,11,2,0,0,1] rd7_xtuple_0
s: 79
d: [2,227,3,4,1,1,44,88,218]
s: 227
d: [2,228,1,11,2,1,44,79,114]
d: [2,238,1,9,1,2,21,228,119]
d: [2,239,1,9,1,3,22,238,104]
d: [2,240,1,9,1,4,22,239,109]
a: [3,244,1,9,1,240]
s: 244
d: [2,261,2,5,2,1,44,244,79]
d: [2,268,2,2,1,2,21,261,99]
d: [2,269,2,0,0,3,21,268,124]
e: 269
d: [2,262,2,5,2,1,44,79,244]
d: [2,271,2,2,1,2,21,262,99]
d: [2,272,2,0,0,3,21,271,124]
e: 272
d: [2,235,1,11,2,1,44,79,104]
d: [2,256,1,9,1,2,21,235,119]
d: [2,257,1,9,1,3,22,256,109]
d: [2,258,1,9,1,4,22,257,114]
f: 272,258,64,65,66,67,68,69
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t4_xtuple_0
% SZS output start Proof for xtuple_0__t4_xtuple_0
1. ! [X0] : (v2_xtuple_0(X0) => ! [X1] : (v2_xtuple_0(X1) => ((k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k4_xtuple_0(X0) = k4_xtuple_0(X1)) => X0 = X1))) [input t4_xtuple_0]
2. ~! [X0] : (v2_xtuple_0(X0) => ! [X1] : (v2_xtuple_0(X1) => ((k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k4_xtuple_0(X0) = k4_xtuple_0(X1)) => X0 = X1))) [negated conjecture 1]
36. ! [X0] : (v2_xtuple_0(X0) => k3_xtuple_0(k4_xtuple_0(X0),k5_xtuple_0(X0),k2_xtuple_0(X0)) = X0) [input rd7_xtuple_0]
46. ? [X0] : (? [X1] : ((X0 != X1 & (k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k4_xtuple_0(X0) = k4_xtuple_0(X1))) & v2_xtuple_0(X1)) & v2_xtuple_0(X0)) [ennf transformation 2]
47. ? [X0] : (? [X1] : (X0 != X1 & k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k4_xtuple_0(X0) = k4_xtuple_0(X1) & v2_xtuple_0(X1)) & v2_xtuple_0(X0)) [flattening 46]
51. ! [X0] : (k3_xtuple_0(k4_xtuple_0(X0),k5_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v2_xtuple_0(X0)) [ennf transformation 36]
53. ? [X0] : (? [X1] : (X0 != X1 & k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k4_xtuple_0(X0) = k4_xtuple_0(X1) & v2_xtuple_0(X1)) & v2_xtuple_0(X0)) => (? [X1] : (sK0 != X1 & k2_xtuple_0(X1) = k2_xtuple_0(sK0) & k5_xtuple_0(X1) = k5_xtuple_0(sK0) & k4_xtuple_0(X1) = k4_xtuple_0(sK0) & v2_xtuple_0(X1)) & v2_xtuple_0(sK0)) [choice axiom]
54. ? [X1] : (sK0 != X1 & k2_xtuple_0(X1) = k2_xtuple_0(sK0) & k5_xtuple_0(X1) = k5_xtuple_0(sK0) & k4_xtuple_0(X1) = k4_xtuple_0(sK0) & v2_xtuple_0(X1)) => (sK0 != sK1 & k2_xtuple_0(sK0) = k2_xtuple_0(sK1) & k5_xtuple_0(sK0) = k5_xtuple_0(sK1) & k4_xtuple_0(sK0) = k4_xtuple_0(sK1) & v2_xtuple_0(sK1)) [choice axiom]
55. (sK0 != sK1 & k2_xtuple_0(sK0) = k2_xtuple_0(sK1) & k5_xtuple_0(sK0) = k5_xtuple_0(sK1) & k4_xtuple_0(sK0) = k4_xtuple_0(sK1) & v2_xtuple_0(sK1)) & v2_xtuple_0(sK0) [skolemisation 47,54,53]
64. v2_xtuple_0(sK0) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. v2_xtuple_0(sK1) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
66. k4_xtuple_0(sK0) = k4_xtuple_0(sK1) [cnf transformation 55] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
67. k5_xtuple_0(sK0) = k5_xtuple_0(sK1) [cnf transformation 55] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
68. k2_xtuple_0(sK0) = k2_xtuple_0(sK1) [cnf transformation 55] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
69. sK0 != sK1 [cnf transformation 55] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
79. k3_xtuple_0(k4_xtuple_0(X0),k5_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v2_xtuple_0(X0) [cnf transformation 51] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
97. 1 <=> sK0 = sK1 [avatar definition]
99. sK0 != sK1 <- (~1) [avatar component clause 97] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. ~1 [avatar split clause 69,97]
102. 2 <=> k2_xtuple_0(sK0) = k2_xtuple_0(sK1) [avatar definition]
104. k2_xtuple_0(sK0) = k2_xtuple_0(sK1) <- (2) [avatar component clause 102] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 2 [avatar split clause 68,102]
107. 3 <=> k5_xtuple_0(sK0) = k5_xtuple_0(sK1) [avatar definition]
109. k5_xtuple_0(sK0) = k5_xtuple_0(sK1) <- (3) [avatar component clause 107] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. 3 [avatar split clause 67,107]
112. 4 <=> k4_xtuple_0(sK0) = k4_xtuple_0(sK1) [avatar definition]
114. k4_xtuple_0(sK0) = k4_xtuple_0(sK1) <- (4) [avatar component clause 112] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
115. 4 [avatar split clause 66,112]
117. 5 <=> v2_xtuple_0(sK1) [avatar definition]
119. v2_xtuple_0(sK1) <- (5) [avatar component clause 117] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
120. 5 [avatar split clause 65,117]
122. 6 <=> v2_xtuple_0(sK0) [avatar definition]
124. v2_xtuple_0(sK0) <- (6) [avatar component clause 122] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
125. 6 [avatar split clause 64,122]
235. sK1 = k3_xtuple_0(k4_xtuple_0(sK1),k5_xtuple_0(sK1),k2_xtuple_0(sK0)) | ~v2_xtuple_0(sK1) <- (2) [superposition 79,104] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
242. 18 <=> sK1 = k3_xtuple_0(k4_xtuple_0(sK0),k5_xtuple_0(sK0),k2_xtuple_0(sK0)) [avatar definition]
244. sK1 = k3_xtuple_0(k4_xtuple_0(sK0),k5_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (18) [avatar component clause 242] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
256. sK1 = k3_xtuple_0(k4_xtuple_0(sK1),k5_xtuple_0(sK1),k2_xtuple_0(sK0)) <- (2, 5) [subsumption resolution 235,119] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
257. sK1 = k3_xtuple_0(k4_xtuple_0(sK1),k5_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (2, 3, 5) [forward demodulation 256,109] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
258. sK1 = k3_xtuple_0(k4_xtuple_0(sK0),k5_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (2, 3, 4, 5) [forward demodulation 257,114] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
259. 18 | ~2 | ~3 | ~4 | ~5 [avatar split clause 258,117,112,107,102,242]
262. sK0 = sK1 | ~v2_xtuple_0(sK0) <- (18) [superposition 79,244] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
271. ~v2_xtuple_0(sK0) <- (~1, 18) [subsumption resolution 262,99] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
272. $false <- (~1, 6, 18) [subsumption resolution 271,124] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
273. 1 | ~6 | ~18 [avatar contradiction clause 272]
274. $false [avatar sat refutation 100,105,110,115,120,125,259,273] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:103,thDist:-103}
% SZS output end Proof for xtuple_0__t4_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 32
% Generated clauses: 134
% Active clauses: 40
% Passive clauses: 63
% Final active clauses: 39
% Final passive clauses: 23
% Main loop iterations started: 40
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 14
% Fw demodulations: 11
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 16
% Fw demodulations to eq. taut.: 5
% 
% >>> Generating Inferences
% Binary resolution: 5
% Forward superposition: 22
% Backward superposition: 27
% Self superposition: 4
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 13
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

