i: [1,124,0,2,1,1,0,0] t18_tops_1
a: [3,182,0,2,1,124]
s: 182
i: [1,125,0,7,1,1,0,0] t18_tops_1
a: [3,177,0,7,1,125]
s: 177
i: [1,126,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,187,0,2,1,126]
s: 187
i: [1,127,0,2,1,0,0,2] fc1_xboole_0
a: [3,192,0,2,1,127]
s: 192
i: [1,170,0,2,1,0,0,2] rc2_xboole_0
a: [3,202,0,2,1,170]
s: 202
i: [1,171,0,2,1,0,0,2] existence_l1_struct_0
a: [3,207,0,2,1,171]
s: 207
i: [1,172,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,212,0,2,1,172]
s: 212
i: [1,173,0,2,1,0,0,2] rc1_xboole_0
a: [3,217,0,2,1,173]
s: 217
i: [1,129,0,3,1,0,0,2] fc1_subset_1
s: 129
i: [1,157,0,3,1,0,0,2] rc2_subset_1
s: 157
i: [1,160,0,3,1,0,0,255] reflexivity_r1_tarski
s: 160
i: [1,128,0,3,1,0,0,255] d2_xboole_0
a: [3,197,0,3,1,128]
s: 197
r: 192
i: [1,141,0,4,2,0,0,2] dt_l1_pre_topc
s: 141
d: [2,220,1,2,1,1,40,141,182]
a: [3,225,1,2,1,220]
s: 225
i: [1,155,0,4,1,0,0,2] existence_m1_subset_1
s: 155
i: [1,159,0,4,1,0,0,255] rc3_subset_1
s: 159
d: [2,221,1,2,1,1,40,141,212]
a: [3,230,1,2,1,221]
s: 230
i: [1,130,0,5,1,0,0,255] t4_boole
s: 130
i: [1,131,0,5,1,0,0,255] t3_boole
s: 131
i: [1,135,0,5,2,0,0,2] rc1_subset_1
s: 135
i: [1,138,0,5,2,0,0,1] fc3_struct_0
s: 138
i: [1,144,0,5,2,0,0,2] t6_boole
s: 144
i: [1,152,0,5,2,0,0,255] t7_boole
s: 152
i: [1,156,0,5,1,0,0,2] rc2_subset_1
s: 156
i: [1,158,0,5,1,0,0,255] rc3_subset_1
s: 158
d: [2,232,0,5,1,1,22,130,197]
s: 232
d: [2,233,0,5,1,1,22,131,197]
s: 233
d: [2,234,0,5,2,1,22,144,197]
s: 234
i: [1,137,0,6,2,0,0,255] rc4_subset_1
s: 137
i: [1,139,0,6,2,0,0,1] d2_struct_0
s: 139
i: [1,147,0,6,2,0,0,255] t1_subset
s: 147
i: [1,161,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 161
d: [2,239,0,6,2,1,22,139,197]
s: 239
d: [2,240,1,7,2,2,44,177,239]
d: [2,243,1,5,1,3,21,240,225]
a: [3,247,1,5,1,243]
s: 247
i: [1,134,0,7,2,0,0,2] rc1_subset_1
s: 134
i: [1,136,0,7,2,0,0,255] rc4_subset_1
s: 136
d: [2,249,2,7,2,2,44,247,234]
s: 249
i: [1,142,0,7,2,0,0,1] fc8_tops_1
s: 142
d: [2,250,1,8,3,1,44,142,239]
d: [2,251,1,6,2,1,21,250,141]
s: 251
i: [1,150,0,7,2,0,0,255] t3_subset
s: 150
i: [1,151,0,7,3,0,0,1] t8_boole
s: 151
d: [2,235,1,7,3,1,44,234,234]
s: 235
i: [1,140,0,8,2,0,0,1] dt_k1_struct_0
s: 140
d: [2,236,1,7,2,1,44,233,234]
s: 236
i: [1,145,0,8,3,0,0,2] cc1_subset_1
s: 145
i: [1,148,0,8,3,0,0,255] t2_subset
s: 148
d: [2,261,1,5,2,0,40,145,158]
s: 261
i: [1,146,0,9,3,0,0,255] cc4_subset_1
s: 146
d: [2,267,1,8,2,0,40,148,156]
d: [2,273,1,5,1,0,21,267,129]
s: 273
i: [1,154,0,9,3,0,0,255] t5_subset
s: 154
i: [1,153,0,10,3,0,0,255] t4_subset
s: 153
d: [2,268,1,8,2,0,40,148,158]
d: [2,274,1,5,1,0,21,268,129]
s: 274
i: [1,162,0,10,2,0,0,255] dt_k3_subset_1
s: 162
d: [2,282,2,5,1,0,40,273,161]
s: 282
i: [1,132,0,11,4,0,0,255] cc3_subset_1
s: 132
i: [1,133,0,11,4,0,0,255] cc2_subset_1
s: 133
d: [2,315,1,9,3,0,40,132,158]
d: [2,320,1,5,2,0,21,315,159]
s: 320
i: [1,149,0,11,2,0,0,255] d4_subset_1
s: 149
d: [2,299,2,5,1,0,40,274,161]
s: 299
i: [1,163,0,11,2,0,0,255] involutiveness_k3_subset_1
s: 163
i: [1,164,0,14,3,0,0,255] dt_k2_pre_topc
s: 164
d: [2,259,1,6,2,0,40,145,155]
s: 259
i: [1,165,0,14,3,0,0,1] dt_k1_tops_1
s: 165
d: [2,266,1,6,2,0,40,148,155]
s: 266
i: [1,166,0,14,3,0,0,1] dt_k1_tops_1
s: 166
i: [1,167,0,16,3,0,0,255] projectivity_k2_pre_topc
s: 167
d: [2,280,1,6,2,0,40,146,156]
s: 280
i: [1,168,0,16,3,0,0,1] projectivity_k1_tops_1
s: 168
d: [2,353,3,18,4,2,44,249,168]
d: [2,356,3,11,3,2,19,353]
d: [2,358,3,9,2,3,21,356,182]
s: 358
i: [1,169,0,16,3,0,0,1] projectivity_k1_tops_1
s: 169
i: [1,143,0,20,3,0,0,255] d1_tops_1
s: 143
d: [2,360,4,7,2,3,40,358,140]
d: [2,369,4,5,1,4,21,360,225]
a: [3,373,4,5,1,369]
s: 373
d: [2,416,5,0,0,2,52,182,373,142]
e: 416
d: [2,417,5,2,1,1,40,373,142]
d: [2,420,5,0,0,2,21,417,182]
e: 420
f: 420,369,243,220,128,124,125
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tops_1__t18_tops_1
% SZS output start Proof for tops_1__t18_tops_1
1. ! [X0] : (l1_pre_topc(X0) => k1_struct_0(X0) = k1_tops_1(X0,k1_struct_0(X0))) [input t18_tops_1]
2. ~! [X0] : (l1_pre_topc(X0) => k1_struct_0(X0) = k1_tops_1(X0,k1_struct_0(X0))) [negated conjecture 1]
38. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
41. o_0_0_xboole_0 = k1_xboole_0 [input d2_xboole_0]
43. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) & l1_pre_topc(X0)) => k1_tops_1(X0,X1) = k1_tops_1(X0,k1_tops_1(X0,X1))) [input projectivity_k1_tops_1]
46. ! [X0] : (l1_struct_0(X0) => m1_subset_1(k1_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0)))) [input dt_k1_struct_0]
49. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
50. ! [X0] : (l1_pre_topc(X0) => v1_xboole_0(k1_tops_1(X0,k1_struct_0(X0)))) [input fc8_tops_1]
51. ! [X0] : (l1_struct_0(X0) => k1_struct_0(X0) = k1_xboole_0) [input d2_struct_0]
63. ? [X0] : (k1_struct_0(X0) != k1_tops_1(X0,k1_struct_0(X0)) & l1_pre_topc(X0)) [ennf transformation 2]
71. ! [X0] : (k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0)) [ennf transformation 51]
72. ! [X0] : (m1_subset_1(k1_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0)) [ennf transformation 46]
73. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 49]
74. ! [X0] : (v1_xboole_0(k1_tops_1(X0,k1_struct_0(X0))) | ~l1_pre_topc(X0)) [ennf transformation 50]
76. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 38]
100. ! [X0,X1] : (k1_tops_1(X0,X1) = k1_tops_1(X0,k1_tops_1(X0,X1)) | (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0))) [ennf transformation 43]
101. ! [X0,X1] : (k1_tops_1(X0,X1) = k1_tops_1(X0,k1_tops_1(X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0)) [flattening 100]
104. ? [X0] : (k1_struct_0(X0) != k1_tops_1(X0,k1_struct_0(X0)) & l1_pre_topc(X0)) => (k1_struct_0(sK0) != k1_tops_1(sK0,k1_struct_0(sK0)) & l1_pre_topc(sK0)) [choice axiom]
105. k1_struct_0(sK0) != k1_tops_1(sK0,k1_struct_0(sK0)) & l1_pre_topc(sK0) [skolemisation 63,104]
124. l1_pre_topc(sK0) [cnf transformation 105] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
125. k1_struct_0(sK0) != k1_tops_1(sK0,k1_struct_0(sK0)) [cnf transformation 105] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
128. o_0_0_xboole_0 = k1_xboole_0 [cnf transformation 41] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
139. k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0) [cnf transformation 71] {a:0,w:6,wCS:2400,sine:1,thAx:0,allAx:1,thDist:-1}
140. m1_subset_1(k1_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) [cnf transformation 72] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
141. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 73] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
142. v1_xboole_0(k1_tops_1(X0,k1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 74] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
144. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 76] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
168. k1_tops_1(X0,X1) = k1_tops_1(X0,k1_tops_1(X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 101] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
175. 1 <=> k1_struct_0(sK0) = k1_tops_1(sK0,k1_struct_0(sK0)) [avatar definition]
177. k1_struct_0(sK0) != k1_tops_1(sK0,k1_struct_0(sK0)) <- (~1) [avatar component clause 175] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
178. ~1 [avatar split clause 125,175]
180. 2 <=> l1_pre_topc(sK0) [avatar definition]
182. l1_pre_topc(sK0) <- (2) [avatar component clause 180] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
183. 2 [avatar split clause 124,180]
195. 5 <=> o_0_0_xboole_0 = k1_xboole_0 [avatar definition]
197. o_0_0_xboole_0 = k1_xboole_0 <- (5) [avatar component clause 195] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
198. 5 [avatar split clause 128,195]
220. l1_struct_0(sK0) <- (2) [resolution 141,182] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
223. 10 <=> l1_struct_0(sK0) [avatar definition]
225. l1_struct_0(sK0) <- (10) [avatar component clause 223] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
226. 10 | ~2 [avatar split clause 220,180,223]
234. o_0_0_xboole_0 = X0 | ~v1_xboole_0(X0) <- (5) [forward demodulation 144,197] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
239. k1_struct_0(X0) = o_0_0_xboole_0 | ~l1_struct_0(X0) <- (5) [forward demodulation 139,197] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
240. o_0_0_xboole_0 != k1_tops_1(sK0,o_0_0_xboole_0) | ~l1_struct_0(sK0) <- (~1, 5) [superposition 177,239] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
243. o_0_0_xboole_0 != k1_tops_1(sK0,o_0_0_xboole_0) <- (~1, 5, 10) [subsumption resolution 240,225] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
245. 12 <=> o_0_0_xboole_0 = k1_tops_1(sK0,o_0_0_xboole_0) [avatar definition]
247. o_0_0_xboole_0 != k1_tops_1(sK0,o_0_0_xboole_0) <- (~12) [avatar component clause 245] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
248. ~12 | 1 | ~5 | ~10 [avatar split clause 243,223,195,175,245]
249. k1_tops_1(sK0,X0) != X0 | ~v1_xboole_0(X0) <- (5, ~12) [superposition 247,234] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
353. k1_tops_1(sK0,X0) != k1_tops_1(sK0,X0) | ~v1_xboole_0(k1_tops_1(sK0,X0)) | ~m1_subset_1(X0,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (5, ~12) [superposition 249,168] {a:3,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
356. ~v1_xboole_0(k1_tops_1(sK0,X0)) | ~m1_subset_1(X0,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (5, ~12) [trivial inequality removal 353] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
358. ~m1_subset_1(X0,k1_zfmisc_1(u1_struct_0(sK0))) | ~v1_xboole_0(k1_tops_1(sK0,X0)) <- (2, 5, ~12) [subsumption resolution 356,182] {a:3,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
360. ~v1_xboole_0(k1_tops_1(sK0,k1_struct_0(sK0))) | ~l1_struct_0(sK0) <- (2, 5, ~12) [resolution 358,140] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
369. ~v1_xboole_0(k1_tops_1(sK0,k1_struct_0(sK0))) <- (2, 5, 10, ~12) [subsumption resolution 360,225] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
371. 13 <=> v1_xboole_0(k1_tops_1(sK0,k1_struct_0(sK0))) [avatar definition]
373. ~v1_xboole_0(k1_tops_1(sK0,k1_struct_0(sK0))) <- (~13) [avatar component clause 371] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
374. ~13 | ~2 | ~5 | ~10 | 12 [avatar split clause 369,245,223,195,180,371]
417. ~l1_pre_topc(sK0) <- (~13) [resolution 373,142] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
420. $false <- (2, ~13) [subsumption resolution 417,182] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
421. ~2 | 13 [avatar contradiction clause 420]
423. $false [avatar sat refutation 178,183,198,226,248,374,421] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:125,thDist:-125}
% SZS output end Proof for tops_1__t18_tops_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 51
% >>> Preprocessing
% Introduced skolems: 10
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 50
% Generated clauses: 221
% Active clauses: 64
% Passive clauses: 137
% Final active clauses: 63
% Final passive clauses: 65
% Main loop iterations started: 72
% 
% >>> Simplifying Inferences
% Duplicate literals: 12
% Trivial inequalities: 1
% Fw subsumption resolutions: 23
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 12
% Forward subsumptions: 19
% 
% >>> Generating Inferences
% Binary resolution: 83
% Unit resulting resolution: 1
% Forward superposition: 10
% Backward superposition: 14
% Self superposition: 4
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 22
% SAT solver unit clauses: 9
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9594
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% backward demodulation index maintenance: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

