i: [1,17,0,2,1,1,0,0] t21_orders_1
a: [3,28,0,2,1,17]
s: 28
i: [1,16,0,2,1,1,0,0] t21_orders_1
a: [3,33,0,2,1,16]
s: 33
i: [1,15,0,2,1,1,0,0] t21_orders_1
a: [3,38,0,2,1,15]
s: 38
i: [1,18,0,6,3,0,0,1] d3_orders_1
s: 18
i: [1,19,0,6,3,0,0,1] d3_orders_1
s: 19
i: [1,21,0,6,3,0,0,255] d4_orders_1
s: 21
d: [2,40,1,4,2,1,40,21,33]
d: [2,41,1,2,1,2,21,40,38]
a: [3,45,1,2,1,41]
s: 45
i: [1,22,0,6,3,0,0,255] d4_orders_1
s: 22
d: [2,47,1,4,2,1,40,22,33]
d: [2,48,1,2,1,2,21,47,38]
a: [3,52,1,2,1,48]
s: 52
i: [1,23,0,6,3,0,0,255] d4_orders_1
s: 23
i: [1,20,0,8,4,0,0,1] d3_orders_1
s: 20
d: [2,61,2,0,0,4,52,38,45,28,52,20]
e: 61
d: [2,62,2,6,3,1,40,20,52]
d: [2,64,2,4,2,2,21,62,28]
d: [2,65,2,2,1,3,21,64,45]
d: [2,66,2,0,0,4,21,65,38]
e: 66
f: 66,48,41,15,16,17
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t21_orders_1
% SZS output start Proof for orders_1__t21_orders_1
1. ! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) => v1_orders_1(X0))) [input t21_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) => v1_orders_1(X0))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => (v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0)))) [input d3_orders_1]
4. ! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d4_orders_1]
5. ? [X0] : ((~v1_orders_1(X0) & v2_orders_1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
6. ? [X0] : (~v1_orders_1(X0) & v2_orders_1(X0) & v1_relat_1(X0)) [flattening 5]
7. ! [X0] : ((v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
8. ! [X0] : ((v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
9. ? [X0] : (~v1_orders_1(X0) & v2_orders_1(X0) & v1_relat_1(X0)) => (~v1_orders_1(sK0) & v2_orders_1(sK0) & v1_relat_1(sK0)) [choice axiom]
10. ~v1_orders_1(sK0) & v2_orders_1(sK0) & v1_relat_1(sK0) [skolemisation 6,9]
11. ! [X0] : (((v1_orders_1(X0) | (~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 7]
12. ! [X0] : (((v1_orders_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 11]
13. ! [X0] : (((v2_orders_1(X0) | (~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 8]
14. ! [X0] : (((v2_orders_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 13]
15. v1_relat_1(sK0) [cnf transformation 10] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
16. v2_orders_1(sK0) [cnf transformation 10] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
17. ~v1_orders_1(sK0) [cnf transformation 10] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
20. ~v8_relat_2(X0) | v1_orders_1(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 12] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
21. ~v2_orders_1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 14] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
22. ~v2_orders_1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 14] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
26. 1 <=> v1_orders_1(sK0) [avatar definition]
28. ~v1_orders_1(sK0) <- (~1) [avatar component clause 26] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
29. ~1 [avatar split clause 17,26]
31. 2 <=> v2_orders_1(sK0) [avatar definition]
33. v2_orders_1(sK0) <- (2) [avatar component clause 31] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
34. 2 [avatar split clause 16,31]
36. 3 <=> v1_relat_1(sK0) [avatar definition]
38. v1_relat_1(sK0) <- (3) [avatar component clause 36] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
39. 3 [avatar split clause 15,36]
40. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 21,33] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
41. v1_relat_2(sK0) <- (2, 3) [subsumption resolution 40,38] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
43. 4 <=> v1_relat_2(sK0) [avatar definition]
45. v1_relat_2(sK0) <- (4) [avatar component clause 43] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
46. 4 | ~2 | ~3 [avatar split clause 41,36,31,43]
47. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 22,33] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
48. v8_relat_2(sK0) <- (2, 3) [subsumption resolution 47,38] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
50. 5 <=> v8_relat_2(sK0) [avatar definition]
52. v8_relat_2(sK0) <- (5) [avatar component clause 50] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
53. 5 | ~2 | ~3 [avatar split clause 48,36,31,50]
62. v1_orders_1(sK0) | ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (5) [resolution 20,52] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
64. ~v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (~1, 5) [subsumption resolution 62,28] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
65. ~v1_relat_1(sK0) <- (~1, 4, 5) [subsumption resolution 64,45] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66. $false <- (~1, 3, 4, 5) [subsumption resolution 65,38] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
67. 1 | ~3 | ~4 | ~5 [avatar contradiction clause 66]
68. $false [avatar sat refutation 29,34,39,46,53,67] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:42,thDist:-42}
% SZS output end Proof for orders_1__t21_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 3
% Introduced skolems: 1
% 
% >>> Saturation
% Initial clauses: 10
% Generated clauses: 27
% Active clauses: 11
% Passive clauses: 13
% Final active clauses: 11
% Final passive clauses: 2
% Main loop iterations started: 11
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 6
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% backward subsumption: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

