i: [1,149,0,2,1,1,0,0] t8_algstr_2
a: [3,219,0,2,1,149]
s: 219
i: [1,148,0,2,1,1,0,0] t8_algstr_2
a: [3,224,0,2,1,148]
s: 224
i: [1,147,0,2,1,1,0,0] t8_algstr_2
a: [3,229,0,2,1,147]
s: 229
i: [1,146,0,2,1,1,0,0] t8_algstr_2
a: [3,234,0,2,1,146]
s: 234
i: [1,145,0,2,1,1,0,0] t8_algstr_2
a: [3,239,0,2,1,145]
s: 239
i: [1,144,0,2,1,1,0,0] t8_algstr_2
a: [3,244,0,2,1,144]
s: 244
i: [1,143,0,2,1,1,0,0] t8_algstr_2
a: [3,249,0,2,1,143]
s: 249
i: [1,142,0,2,1,1,0,0] t8_algstr_2
a: [3,254,0,2,1,142]
s: 254
i: [1,141,0,2,1,1,0,0] t8_algstr_2
a: [3,259,0,2,1,141]
s: 259
i: [1,197,0,2,1,0,0,255] existence_l4_struct_0
a: [3,264,0,2,1,197]
s: 264
i: [1,198,0,2,1,0,0,1] existence_l6_algstr_0
a: [3,269,0,2,1,198]
s: 269
i: [1,199,0,2,1,0,0,2] existence_l3_algstr_0
a: [3,274,0,2,1,199]
s: 274
i: [1,200,0,2,1,0,0,2] existence_l2_algstr_0
a: [3,279,0,2,1,200]
s: 279
i: [1,201,0,2,1,0,0,4] existence_l1_algstr_0
a: [3,284,0,2,1,201]
s: 284
i: [1,202,0,2,1,0,0,2] existence_l3_struct_0
a: [3,289,0,2,1,202]
s: 289
i: [1,203,0,2,1,0,0,2] existence_l5_algstr_0
a: [3,294,0,2,1,203]
s: 294
i: [1,204,0,2,1,0,0,2] existence_l4_algstr_0
a: [3,299,0,2,1,204]
s: 299
i: [1,207,0,2,1,0,0,1] rc2_vectsp_1
a: [3,304,0,2,1,207]
s: 304
i: [1,206,0,2,1,0,0,1] rc2_vectsp_1
a: [3,309,0,2,1,206]
s: 309
i: [1,205,0,2,1,0,0,1] rc2_vectsp_1
a: [3,314,0,2,1,205]
s: 314
i: [1,210,0,2,1,0,0,1] rc11_vectsp_1
a: [3,319,0,2,1,210]
s: 319
i: [1,209,0,2,1,0,0,1] rc11_vectsp_1
a: [3,324,0,2,1,209]
s: 324
i: [1,208,0,2,1,0,0,1] rc11_vectsp_1
a: [3,329,0,2,1,208]
s: 329
i: [1,150,0,13,1,1,0,0] t8_algstr_2
a: [3,214,0,13,1,150]
s: 214
i: [1,151,0,4,2,0,0,255] dt_l4_struct_0
s: 151
i: [1,152,0,4,2,0,0,1] dt_l6_algstr_0
s: 152
d: [2,337,1,2,1,1,40,152,219]
a: [3,342,1,2,1,337]
s: 342
i: [1,153,0,4,2,0,0,1] dt_l6_algstr_0
s: 153
d: [2,349,1,2,1,1,40,153,219]
a: [3,354,1,2,1,349]
s: 354
i: [1,155,0,4,2,0,0,255] dt_l2_algstr_0
s: 155
i: [1,172,0,4,2,0,0,255] dt_l5_algstr_0
s: 172
d: [2,362,2,2,1,1,40,155,342]
a: [3,371,2,2,1,362]
s: 371
i: [1,173,0,4,2,0,0,255] dt_l5_algstr_0
s: 173
d: [2,375,2,2,1,1,40,172,354]
a: [3,389,2,2,1,375]
s: 389
i: [1,179,0,4,2,0,0,2] dt_l4_algstr_0
s: 179
i: [1,180,0,4,2,0,0,2] dt_l4_algstr_0
s: 180
d: [2,393,2,2,1,1,40,173,354]
a: [3,407,2,2,1,393]
s: 407
i: [1,194,0,4,1,0,0,2] existence_m1_subset_1
s: 194
d: [2,411,3,2,1,1,40,179,389]
a: [3,425,3,2,1,411]
s: 425
i: [1,165,0,6,3,0,0,255] cc4_algstr_0
s: 165
i: [1,166,0,6,3,0,0,255] cc4_algstr_0
s: 166
d: [2,429,3,2,1,1,40,180,389]
a: [3,443,3,2,1,429]
s: 443
i: [1,174,0,6,3,0,0,255] cc26_algstr_0
s: 174
d: [2,331,1,2,1,1,40,151,264]
a: [3,335,1,2,1,331]
s: 335
i: [1,175,0,6,3,0,0,255] cc26_algstr_0
s: 175
i: [1,154,0,7,2,0,0,255] dt_k1_group_1
s: 154
d: [2,338,1,2,1,1,40,152,269]
a: [3,347,1,2,1,338]
s: 347
i: [1,171,0,7,2,0,0,1] dt_k5_struct_0
s: 171
d: [2,350,1,2,1,1,40,153,269]
a: [3,359,1,2,1,350]
s: 359
i: [1,157,0,8,4,0,0,255] cc3_algstr_1
s: 157
i: [1,159,0,8,4,0,0,255] cc1_algstr_1
s: 159
d: [2,466,2,6,3,1,40,157,342]
d: [2,481,2,4,2,2,21,466,249]
d: [2,482,2,2,1,3,21,481,259]
a: [3,486,2,2,1,482]
s: 486
i: [1,160,0,8,4,0,0,255] cc1_algstr_1
s: 160
d: [2,502,2,6,3,1,40,159,342]
d: [2,509,2,4,2,2,21,502,249]
d: [2,510,2,2,1,3,21,509,259]
a: [3,514,2,2,1,510]
s: 514
i: [1,161,0,8,4,0,0,255] cc1_algstr_1
s: 161
i: [1,162,0,8,4,0,0,255] cc1_algstr_1
s: 162
d: [2,522,2,6,3,1,40,160,342]
d: [2,529,2,4,2,2,21,522,249]
d: [2,530,2,2,1,3,21,529,259]
a: [3,534,2,2,1,530]
s: 534
i: [1,167,0,8,4,0,0,255] cc3_algstr_0
s: 167
d: [2,542,2,6,3,1,40,162,342]
d: [2,549,2,4,2,2,21,542,249]
d: [2,550,2,2,1,3,21,549,259]
a: [3,554,2,2,1,550]
s: 554
i: [1,177,0,8,4,0,0,255] cc6_algstr_1
s: 177
i: [1,178,0,8,4,0,0,255] cc25_algstr_0
s: 178
d: [2,569,2,6,3,1,40,177,354]
d: [2,573,2,4,2,2,21,569,244]
d: [2,574,2,2,1,3,21,573,259]
a: [3,578,2,2,1,574]
s: 578
i: [1,183,0,8,4,0,0,255] cc3_vectsp_1
s: 183
d: [2,561,3,6,3,1,40,167,534]
d: [2,562,3,4,2,2,21,561,514]
d: [2,563,3,2,1,3,21,562,371]
a: [3,567,3,2,1,563]
s: 567
i: [1,184,0,8,4,0,0,255] cc3_vectsp_1
s: 184
i: [1,188,0,11,4,0,0,255] ie6_vectsp_1
s: 188
d: [2,612,3,4,2,1,40,578,175]
d: [2,614,3,2,1,2,21,612,354]
a: [3,618,3,2,1,614]
s: 618
i: [1,190,0,11,4,0,0,255] d8_vectsp_1
s: 190
d: [2,613,3,4,2,1,40,578,174]
d: [2,620,3,2,1,2,21,613,354]
a: [3,624,3,2,1,620]
s: 624
i: [1,192,0,11,4,0,0,255] cc6_algstr_0
s: 192
i: [1,193,0,11,4,0,0,255] cc5_algstr_0
s: 193
d: [2,628,3,6,3,1,40,183,389]
d: [2,649,3,4,2,2,21,628,224]
d: [2,650,3,2,1,3,21,649,259]
a: [3,654,3,2,1,650]
s: 654
i: [1,168,0,12,4,0,0,255] cc2_algstr_0
s: 168
d: [2,660,3,6,3,1,40,184,389]
d: [2,673,3,4,2,2,21,660,224]
d: [2,674,3,2,1,3,21,673,259]
a: [3,678,3,2,1,674]
s: 678
i: [1,169,0,12,4,0,0,255] cc2_algstr_0
s: 169
i: [1,164,0,14,7,0,0,255] cc2_algstr_1
s: 164
d: [2,361,1,2,1,1,40,155,279]
a: [3,366,1,2,1,361]
s: 366
i: [1,170,0,15,5,0,0,255] cc1_algstr_0
s: 170
d: [2,373,1,2,1,1,40,172,294]
a: [3,379,1,2,1,373]
s: 379
i: [1,191,0,15,4,0,0,255] d8_vectsp_1
s: 191
i: [1,189,0,17,5,0,0,255] d8_vectsp_1
s: 189
d: [2,374,1,2,1,1,40,172,314]
a: [3,384,1,2,1,374]
s: 384
i: [1,196,0,17,4,0,0,1] dt_k6_algstr_0
s: 196
d: [2,391,1,2,1,1,40,173,294]
a: [3,397,1,2,1,391]
s: 397
i: [1,195,0,18,6,0,0,1] dt_k1_algstr_2
s: 195
i: [1,187,0,21,7,0,0,1] t3_algstr_2
s: 187
d: [2,392,1,2,1,1,40,173,314]
a: [3,402,1,2,1,392]
s: 402
i: [1,186,0,37,11,0,0,1] t7_algstr_2
s: 186
d: [2,787,1,41,11,1,44,214,186]
d: [2,794,1,39,10,2,21,787,219]
d: [2,795,1,37,9,3,21,794,224]
d: [2,796,1,35,8,4,21,795,229]
d: [2,797,1,33,7,5,21,796,234]
d: [2,798,1,31,6,6,21,797,244]
d: [2,799,1,29,5,7,21,798,249]
d: [2,800,1,27,4,8,21,799,254]
d: [2,801,1,25,3,9,21,800,259]
a: [3,805,1,5,1,801]
s: 805
d: [2,815,4,0,0,2,52,425,805,171]
e: 815
d: [2,816,2,2,1,1,40,805,171]
d: [2,818,2,0,0,2,21,816,425]
e: 818
r: 805
a: [3,804,1,5,1,801]
s: 804
a: [3,809,1,7,1,801]
s: 809
d: [2,849,3,0,0,6,52,259,514,554,342,804,809,195]
e: 849
d: [2,850,2,13,5,1,40,809,195]
d: [2,852,2,8,4,2,21,850,804]
d: [2,853,2,6,3,3,21,852,342]
d: [2,854,2,4,2,4,21,853,554]
d: [2,855,2,2,1,5,21,854,514]
d: [2,856,2,0,0,6,21,855,259]
e: 856
r: 809
d: [2,820,2,14,4,1,40,804,170]
d: [2,825,2,12,3,2,21,820,371]
a: [3,829,2,4,1,825]
s: 829
a: [3,808,1,7,1,801]
s: 808
a: [3,832,2,4,1,825]
s: 832
d: [2,409,1,2,1,1,40,179,299]
a: [3,415,1,2,1,409]
s: 415
d: [2,410,1,2,1,1,40,179,329]
a: [3,420,1,2,1,410]
s: 420
a: [3,836,2,4,1,825]
s: 836
d: [2,427,1,2,1,1,40,180,299]
a: [3,433,1,2,1,427]
s: 433
d: [2,858,2,20,4,1,40,808,170]
d: [2,863,2,18,3,2,21,858,371]
a: [3,867,2,6,1,863]
s: 867
d: [2,428,1,2,1,1,40,180,329]
a: [3,438,1,2,1,428]
s: 438
d: [2,465,1,6,3,1,40,157,279]
a: [3,479,1,2,1,465]
s: 479
a: [3,870,2,6,1,863]
s: 870
d: [2,571,1,6,3,1,40,177,294]
a: [3,600,1,2,1,571]
s: 600
a: [3,874,2,6,1,863]
s: 874
d: [2,572,1,6,3,1,40,177,314]
d: [2,602,1,4,2,2,21,572,309]
a: [3,610,1,2,1,602]
s: 610
d: [2,626,1,6,3,1,40,183,299]
a: [3,640,1,2,1,626]
s: 640
d: [2,627,1,6,3,1,40,183,329]
d: [2,642,1,4,2,2,21,627,319]
d: [2,643,1,2,1,3,21,642,324]
a: [3,647,1,2,1,643]
s: 647
d: [2,659,1,6,3,1,40,184,329]
d: [2,666,1,4,2,2,21,659,319]
d: [2,667,1,2,1,3,21,666,324]
a: [3,671,1,2,1,667]
s: 671
d: [2,447,2,2,1,1,40,347,155]
a: [3,451,2,2,1,447]
s: 451
a: [3,813,1,13,1,801]
s: 813
d: [2,686,1,9,3,0,40,192,194]
s: 686
d: [2,453,2,2,1,1,40,359,173]
a: [3,458,2,2,1,453]
s: 458
d: [2,690,1,9,3,0,40,193,194]
s: 690
d: [2,454,2,2,1,1,40,359,172]
a: [3,463,2,2,1,454]
s: 463
d: [2,683,1,10,4,0,40,192,171]
s: 683
d: [2,684,1,10,4,0,40,192,154]
s: 684
d: [2,467,2,6,3,1,40,157,347]
a: [3,499,2,2,1,467]
s: 499
d: [2,687,1,10,4,0,40,193,171]
s: 687
d: [2,570,2,6,3,1,40,177,359]
a: [3,587,2,2,1,570]
s: 587
d: [2,688,1,10,4,0,40,193,154]
s: 688
d: [2,680,1,13,5,0,44,154,188]
d: [2,681,1,11,4,0,21,680,180]
s: 681
d: [2,704,2,6,3,1,40,379,184]
a: [3,715,2,2,1,704]
s: 715
d: [2,691,1,12,4,0,40,168,171]
s: 691
d: [2,706,2,2,1,1,40,379,180]
a: [3,725,2,2,1,706]
s: 725
d: [2,692,1,12,4,0,40,168,154]
s: 692
d: [2,694,1,12,3,0,40,168,194]
s: 694
d: [2,707,2,2,1,1,40,379,179]
a: [3,730,2,2,1,707]
s: 730
d: [2,695,1,12,4,0,40,169,171]
s: 695
d: [2,734,2,6,3,1,40,384,184]
d: [2,738,2,4,2,2,21,734,304]
d: [2,739,2,2,1,3,21,738,309]
a: [3,743,2,2,1,739]
s: 743
d: [2,696,1,12,4,0,40,169,154]
s: 696
d: [2,698,1,12,3,0,40,169,194]
s: 698
d: [2,735,2,6,3,1,40,384,183]
d: [2,745,2,4,2,2,21,735,304]
d: [2,746,2,2,1,3,21,745,309]
a: [3,750,2,2,1,746]
s: 750
d: [2,685,1,14,6,0,40,192,190]
s: 685
d: [2,736,2,2,1,1,40,384,180]
a: [3,755,2,2,1,736]
s: 755
d: [2,689,1,14,6,0,40,193,190]
s: 689
d: [2,693,1,16,6,0,40,168,190]
s: 693
d: [2,737,2,2,1,1,40,384,179]
a: [3,760,2,2,1,737]
s: 760
d: [2,697,1,16,6,0,40,169,190]
s: 697
d: [2,897,3,2,1,1,40,458,151]
a: [3,901,3,2,1,897]
s: 901
d: [2,700,1,16,5,0,40,170,171]
s: 700
d: [2,701,1,16,5,0,40,170,154]
s: 701
d: [2,903,3,6,3,1,40,463,184]
a: [3,914,3,2,1,903]
s: 914
d: [2,703,1,17,4,0,40,170,194]
s: 703
d: [2,905,3,2,1,1,40,463,180]
a: [3,924,3,2,1,905]
s: 924
d: [2,774,1,21,8,0,40,195,193]
d: [2,776,1,19,7,0,18,774]
d: [2,780,1,17,6,0,21,776,155]
s: 780
d: [2,775,1,21,8,0,40,195,192]
d: [2,781,1,19,7,0,21,775,155]
s: 781
d: [2,887,2,22,5,1,44,813,189]
d: [2,888,2,15,4,2,21,887,808]
d: [2,889,2,13,3,3,21,888,678]
d: [2,890,2,11,2,4,21,889,389]
d: [2,891,2,9,1,5,21,890,259]
a: [3,895,2,9,1,891]
s: 895
d: [2,957,3,0,0,7,52,259,514,554,239,342,804,895,187]
e: 957
d: [2,958,3,20,7,1,44,895,187]
d: [2,959,3,15,6,1,19,958]
d: [2,961,3,10,5,2,21,959,804]
d: [2,962,3,8,4,3,21,961,342]
d: [2,963,3,6,3,4,21,962,239]
d: [2,964,3,4,2,5,21,963,554]
d: [2,965,3,2,1,6,21,964,514]
d: [2,966,3,0,0,7,21,965,259]
e: 966
d: [2,445,3,2,1,1,40,407,151]
f: 966,891,856,818,801,674,550,510,445,393,375,349,337,141,142,143,144,145,146,147,148,149,150
% Refutation found. Thanks to Tanya!
% SZS status Theorem for algstr_2__t8_algstr_2
% SZS output start Proof for algstr_2__t8_algstr_2
1. ! [X0] : ((l6_algstr_0(X0) & v4_vectsp_1(X0) & v2_vectsp_1(X0) & v4_rlvect_1(X0) & v2_rlvect_1(X0) & v7_algstr_1(X0) & v4_algstr_1(X0) & v1_algstr_1(X0) & ~v2_struct_0(X0)) => k5_struct_0(X0) = k6_algstr_0(X0,k1_algstr_2(X0,k5_struct_0(X0)),k1_algstr_2(X0,k5_struct_0(X0)))) [input t8_algstr_2]
2. ~! [X0] : ((l6_algstr_0(X0) & v4_vectsp_1(X0) & v2_vectsp_1(X0) & v4_rlvect_1(X0) & v2_rlvect_1(X0) & v7_algstr_1(X0) & v4_algstr_1(X0) & v1_algstr_1(X0) & ~v2_struct_0(X0)) => k5_struct_0(X0) = k6_algstr_0(X0,k1_algstr_2(X0,k5_struct_0(X0)),k1_algstr_2(X0,k5_struct_0(X0)))) [negated conjecture 1]
4. ! [X0] : (l2_algstr_0(X0) => ((v4_algstr_1(X0) & ~v2_struct_0(X0)) => (v3_algstr_1(X0) & v2_algstr_1(X0) & v6_algstr_0(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)))) [input cc1_algstr_1]
11. ! [X0] : (l4_algstr_0(X0) => ((v4_vectsp_1(X0) & ~v2_struct_0(X0)) => (v6_vectsp_1(X0) & v3_vectsp_1(X0) & ~v2_struct_0(X0)))) [input cc3_vectsp_1]
18. ! [X0] : ((l4_algstr_0(X0) & ~v2_struct_0(X0)) => (v6_vectsp_1(X0) <=> ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k6_algstr_0(X0,k5_struct_0(X0),X1) = X1))) [input d8_vectsp_1]
19. ! [X0,X1] : ((m1_subset_1(X1,u1_struct_0(X0)) & l2_algstr_0(X0) & v3_algstr_1(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) => m1_subset_1(k1_algstr_2(X0,X1),u1_struct_0(X0))) [input dt_k1_algstr_2]
21. ! [X0] : (l3_struct_0(X0) => m1_subset_1(k5_struct_0(X0),u1_struct_0(X0))) [input dt_k5_struct_0]
30. ! [X0] : (l4_struct_0(X0) => (l3_struct_0(X0) & l2_struct_0(X0))) [input dt_l4_struct_0]
31. ! [X0] : (l5_algstr_0(X0) => (l4_struct_0(X0) & l4_algstr_0(X0))) [input dt_l5_algstr_0]
32. ! [X0] : (l6_algstr_0(X0) => (l5_algstr_0(X0) & l2_algstr_0(X0))) [input dt_l6_algstr_0]
49. ! [X0] : ((l2_algstr_0(X0) & v2_rlvect_1(X0) & v3_algstr_1(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k1_algstr_2(X0,k1_algstr_2(X0,X1)) = X1)) [input t3_algstr_2]
50. ! [X0] : ((l6_algstr_0(X0) & v4_vectsp_1(X0) & v2_vectsp_1(X0) & v4_rlvect_1(X0) & v7_algstr_1(X0) & v4_algstr_1(X0) & v1_algstr_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => k6_algstr_0(X0,k1_algstr_2(X0,X1),X2) = k1_algstr_2(X0,k6_algstr_0(X0,X1,X2))))) [input t7_algstr_2]
56. ! [X0] : (l4_struct_0(X0) => l3_struct_0(X0)) [pure predicate removal 30]
62. ? [X0] : (k5_struct_0(X0) != k6_algstr_0(X0,k1_algstr_2(X0,k5_struct_0(X0)),k1_algstr_2(X0,k5_struct_0(X0))) & (l6_algstr_0(X0) & v4_vectsp_1(X0) & v2_vectsp_1(X0) & v4_rlvect_1(X0) & v2_rlvect_1(X0) & v7_algstr_1(X0) & v4_algstr_1(X0) & v1_algstr_1(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
63. ? [X0] : (k5_struct_0(X0) != k6_algstr_0(X0,k1_algstr_2(X0,k5_struct_0(X0)),k1_algstr_2(X0,k5_struct_0(X0))) & l6_algstr_0(X0) & v4_vectsp_1(X0) & v2_vectsp_1(X0) & v4_rlvect_1(X0) & v2_rlvect_1(X0) & v7_algstr_1(X0) & v4_algstr_1(X0) & v1_algstr_1(X0) & ~v2_struct_0(X0)) [flattening 62]
64. ! [X0] : (l3_struct_0(X0) | ~l4_struct_0(X0)) [ennf transformation 56]
65. ! [X0] : ((l5_algstr_0(X0) & l2_algstr_0(X0)) | ~l6_algstr_0(X0)) [ennf transformation 32]
70. ! [X0] : (((v3_algstr_1(X0) & v2_algstr_1(X0) & v6_algstr_0(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) | (~v4_algstr_1(X0) | v2_struct_0(X0))) | ~l2_algstr_0(X0)) [ennf transformation 4]
71. ! [X0] : ((v3_algstr_1(X0) & v2_algstr_1(X0) & v6_algstr_0(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) | ~v4_algstr_1(X0) | v2_struct_0(X0) | ~l2_algstr_0(X0)) [flattening 70]
82. ! [X0] : (m1_subset_1(k5_struct_0(X0),u1_struct_0(X0)) | ~l3_struct_0(X0)) [ennf transformation 21]
83. ! [X0] : ((l4_struct_0(X0) & l4_algstr_0(X0)) | ~l5_algstr_0(X0)) [ennf transformation 31]
93. ! [X0] : (((v6_vectsp_1(X0) & v3_vectsp_1(X0) & ~v2_struct_0(X0)) | (~v4_vectsp_1(X0) | v2_struct_0(X0))) | ~l4_algstr_0(X0)) [ennf transformation 11]
94. ! [X0] : ((v6_vectsp_1(X0) & v3_vectsp_1(X0) & ~v2_struct_0(X0)) | ~v4_vectsp_1(X0) | v2_struct_0(X0) | ~l4_algstr_0(X0)) [flattening 93]
97. ! [X0] : (! [X1] : (! [X2] : (k6_algstr_0(X0,k1_algstr_2(X0,X1),X2) = k1_algstr_2(X0,k6_algstr_0(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l6_algstr_0(X0) | ~v4_vectsp_1(X0) | ~v2_vectsp_1(X0) | ~v4_rlvect_1(X0) | ~v7_algstr_1(X0) | ~v4_algstr_1(X0) | ~v1_algstr_1(X0) | v2_struct_0(X0))) [ennf transformation 50]
98. ! [X0] : (! [X1] : (! [X2] : (k6_algstr_0(X0,k1_algstr_2(X0,X1),X2) = k1_algstr_2(X0,k6_algstr_0(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l6_algstr_0(X0) | ~v4_vectsp_1(X0) | ~v2_vectsp_1(X0) | ~v4_rlvect_1(X0) | ~v7_algstr_1(X0) | ~v4_algstr_1(X0) | ~v1_algstr_1(X0) | v2_struct_0(X0)) [flattening 97]
99. ! [X0] : (! [X1] : (k1_algstr_2(X0,k1_algstr_2(X0,X1)) = X1 | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_algstr_0(X0) | ~v2_rlvect_1(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 49]
100. ! [X0] : (! [X1] : (k1_algstr_2(X0,k1_algstr_2(X0,X1)) = X1 | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_algstr_0(X0) | ~v2_rlvect_1(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0)) [flattening 99]
103. ! [X0] : ((v6_vectsp_1(X0) <=> ! [X1] : (k6_algstr_0(X0,k5_struct_0(X0),X1) = X1 | ~m1_subset_1(X1,u1_struct_0(X0)))) | (~l4_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 18]
104. ! [X0] : ((v6_vectsp_1(X0) <=> ! [X1] : (k6_algstr_0(X0,k5_struct_0(X0),X1) = X1 | ~m1_subset_1(X1,u1_struct_0(X0)))) | ~l4_algstr_0(X0) | v2_struct_0(X0)) [flattening 103]
109. ! [X0,X1] : (m1_subset_1(k1_algstr_2(X0,X1),u1_struct_0(X0)) | (~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 19]
110. ! [X0,X1] : (m1_subset_1(k1_algstr_2(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0)) [flattening 109]
113. ? [X0] : (k5_struct_0(X0) != k6_algstr_0(X0,k1_algstr_2(X0,k5_struct_0(X0)),k1_algstr_2(X0,k5_struct_0(X0))) & l6_algstr_0(X0) & v4_vectsp_1(X0) & v2_vectsp_1(X0) & v4_rlvect_1(X0) & v2_rlvect_1(X0) & v7_algstr_1(X0) & v4_algstr_1(X0) & v1_algstr_1(X0) & ~v2_struct_0(X0)) => (k5_struct_0(sK0) != k6_algstr_0(sK0,k1_algstr_2(sK0,k5_struct_0(sK0)),k1_algstr_2(sK0,k5_struct_0(sK0))) & l6_algstr_0(sK0) & v4_vectsp_1(sK0) & v2_vectsp_1(sK0) & v4_rlvect_1(sK0) & v2_rlvect_1(sK0) & v7_algstr_1(sK0) & v4_algstr_1(sK0) & v1_algstr_1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
114. k5_struct_0(sK0) != k6_algstr_0(sK0,k1_algstr_2(sK0,k5_struct_0(sK0)),k1_algstr_2(sK0,k5_struct_0(sK0))) & l6_algstr_0(sK0) & v4_vectsp_1(sK0) & v2_vectsp_1(sK0) & v4_rlvect_1(sK0) & v2_rlvect_1(sK0) & v7_algstr_1(sK0) & v4_algstr_1(sK0) & v1_algstr_1(sK0) & ~v2_struct_0(sK0) [skolemisation 63,113]
115. ! [X0] : (((v6_vectsp_1(X0) | ? [X1] : (k6_algstr_0(X0,k5_struct_0(X0),X1) != X1 & m1_subset_1(X1,u1_struct_0(X0)))) & (! [X1] : (k6_algstr_0(X0,k5_struct_0(X0),X1) = X1 | ~m1_subset_1(X1,u1_struct_0(X0))) | ~v6_vectsp_1(X0))) | ~l4_algstr_0(X0) | v2_struct_0(X0)) [nnf transformation 104]
116. ! [X0] : (((v6_vectsp_1(X0) | ? [X1] : (k6_algstr_0(X0,k5_struct_0(X0),X1) != X1 & m1_subset_1(X1,u1_struct_0(X0)))) & (! [X2] : (k6_algstr_0(X0,k5_struct_0(X0),X2) = X2 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~v6_vectsp_1(X0))) | ~l4_algstr_0(X0) | v2_struct_0(X0)) [rectify 115]
117. ! [X0] : (? [X1] : (k6_algstr_0(X0,k5_struct_0(X0),X1) != X1 & m1_subset_1(X1,u1_struct_0(X0))) => (sK1(X0) != k6_algstr_0(X0,k5_struct_0(X0),sK1(X0)) & m1_subset_1(sK1(X0),u1_struct_0(X0)))) [choice axiom]
118. ! [X0] : (((v6_vectsp_1(X0) | (sK1(X0) != k6_algstr_0(X0,k5_struct_0(X0),sK1(X0)) & m1_subset_1(sK1(X0),u1_struct_0(X0)))) & (! [X2] : (k6_algstr_0(X0,k5_struct_0(X0),X2) = X2 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~v6_vectsp_1(X0))) | ~l4_algstr_0(X0) | v2_struct_0(X0)) [skolemisation 116,117]
141. ~v2_struct_0(sK0) [cnf transformation 114] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
142. v1_algstr_1(sK0) [cnf transformation 114] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
143. v4_algstr_1(sK0) [cnf transformation 114] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
144. v7_algstr_1(sK0) [cnf transformation 114] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
145. v2_rlvect_1(sK0) [cnf transformation 114] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
146. v4_rlvect_1(sK0) [cnf transformation 114] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
147. v2_vectsp_1(sK0) [cnf transformation 114] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
148. v4_vectsp_1(sK0) [cnf transformation 114] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
149. l6_algstr_0(sK0) [cnf transformation 114] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
150. k5_struct_0(sK0) != k6_algstr_0(sK0,k1_algstr_2(sK0,k5_struct_0(sK0)),k1_algstr_2(sK0,k5_struct_0(sK0))) [cnf transformation 114] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
151. ~l4_struct_0(X0) | l3_struct_0(X0) [cnf transformation 64] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
152. ~l6_algstr_0(X0) | l2_algstr_0(X0) [cnf transformation 65] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
153. ~l6_algstr_0(X0) | l5_algstr_0(X0) [cnf transformation 65] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
159. ~l2_algstr_0(X0) | ~v4_algstr_1(X0) | v2_struct_0(X0) | v5_algstr_0(X0) [cnf transformation 71] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
162. ~l2_algstr_0(X0) | ~v4_algstr_1(X0) | v2_struct_0(X0) | v3_algstr_1(X0) [cnf transformation 71] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
171. m1_subset_1(k5_struct_0(X0),u1_struct_0(X0)) | ~l3_struct_0(X0) [cnf transformation 82] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
172. ~l5_algstr_0(X0) | l4_algstr_0(X0) [cnf transformation 83] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
173. ~l5_algstr_0(X0) | l4_struct_0(X0) [cnf transformation 83] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
184. ~l4_algstr_0(X0) | ~v4_vectsp_1(X0) | v2_struct_0(X0) | v6_vectsp_1(X0) [cnf transformation 94] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
186. k6_algstr_0(X0,k1_algstr_2(X0,X1),X2) = k1_algstr_2(X0,k6_algstr_0(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l6_algstr_0(X0) | ~v4_vectsp_1(X0) | ~v2_vectsp_1(X0) | ~v4_rlvect_1(X0) | ~v7_algstr_1(X0) | ~v4_algstr_1(X0) | ~v1_algstr_1(X0) | v2_struct_0(X0) [cnf transformation 98] {a:0,w:37,wCS:14800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
187. k1_algstr_2(X0,k1_algstr_2(X0,X1)) = X1 | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v2_rlvect_1(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 100] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
189. k6_algstr_0(X0,k5_struct_0(X0),X2) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~v6_vectsp_1(X0) | ~l4_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 118] {a:0,w:17,wCS:6800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
195. m1_subset_1(k1_algstr_2(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 110] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
212. 1 <=> k5_struct_0(sK0) = k6_algstr_0(sK0,k1_algstr_2(sK0,k5_struct_0(sK0)),k1_algstr_2(sK0,k5_struct_0(sK0))) [avatar definition]
214. k5_struct_0(sK0) != k6_algstr_0(sK0,k1_algstr_2(sK0,k5_struct_0(sK0)),k1_algstr_2(sK0,k5_struct_0(sK0))) <- (~1) [avatar component clause 212] {a:0,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
215. ~1 [avatar split clause 150,212]
217. 2 <=> l6_algstr_0(sK0) [avatar definition]
219. l6_algstr_0(sK0) <- (2) [avatar component clause 217] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
220. 2 [avatar split clause 149,217]
222. 3 <=> v4_vectsp_1(sK0) [avatar definition]
224. v4_vectsp_1(sK0) <- (3) [avatar component clause 222] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
225. 3 [avatar split clause 148,222]
227. 4 <=> v2_vectsp_1(sK0) [avatar definition]
229. v2_vectsp_1(sK0) <- (4) [avatar component clause 227] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
230. 4 [avatar split clause 147,227]
232. 5 <=> v4_rlvect_1(sK0) [avatar definition]
234. v4_rlvect_1(sK0) <- (5) [avatar component clause 232] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
235. 5 [avatar split clause 146,232]
237. 6 <=> v2_rlvect_1(sK0) [avatar definition]
239. v2_rlvect_1(sK0) <- (6) [avatar component clause 237] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
240. 6 [avatar split clause 145,237]
242. 7 <=> v7_algstr_1(sK0) [avatar definition]
244. v7_algstr_1(sK0) <- (7) [avatar component clause 242] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
245. 7 [avatar split clause 144,242]
247. 8 <=> v4_algstr_1(sK0) [avatar definition]
249. v4_algstr_1(sK0) <- (8) [avatar component clause 247] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
250. 8 [avatar split clause 143,247]
252. 9 <=> v1_algstr_1(sK0) [avatar definition]
254. v1_algstr_1(sK0) <- (9) [avatar component clause 252] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
255. 9 [avatar split clause 142,252]
257. 10 <=> v2_struct_0(sK0) [avatar definition]
259. ~v2_struct_0(sK0) <- (~10) [avatar component clause 257] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
260. ~10 [avatar split clause 141,257]
337. l2_algstr_0(sK0) <- (2) [resolution 152,219] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
340. 26 <=> l2_algstr_0(sK0) [avatar definition]
342. l2_algstr_0(sK0) <- (26) [avatar component clause 340] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
343. 26 | ~2 [avatar split clause 337,217,340]
349. l5_algstr_0(sK0) <- (2) [resolution 153,219] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
352. 28 <=> l5_algstr_0(sK0) [avatar definition]
354. l5_algstr_0(sK0) <- (28) [avatar component clause 352] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
355. 28 | ~2 [avatar split clause 349,217,352]
375. l4_algstr_0(sK0) <- (28) [resolution 172,354] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
387. 34 <=> l4_algstr_0(sK0) [avatar definition]
389. l4_algstr_0(sK0) <- (34) [avatar component clause 387] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
390. 34 | ~28 [avatar split clause 375,352,387]
393. l4_struct_0(sK0) <- (28) [resolution 173,354] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
405. 37 <=> l4_struct_0(sK0) [avatar definition]
407. l4_struct_0(sK0) <- (37) [avatar component clause 405] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
408. 37 | ~28 [avatar split clause 393,352,405]
423. 40 <=> l3_struct_0(sK0) [avatar definition]
425. l3_struct_0(sK0) <- (40) [avatar component clause 423] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
445. l3_struct_0(sK0) <- (37) [resolution 407,151] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
446. 40 | ~37 [avatar split clause 445,405,423]
502. ~v4_algstr_1(sK0) | v2_struct_0(sK0) | v5_algstr_0(sK0) <- (26) [resolution 159,342] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
509. v2_struct_0(sK0) | v5_algstr_0(sK0) <- (8, 26) [subsumption resolution 502,249] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
510. v5_algstr_0(sK0) <- (8, ~10, 26) [subsumption resolution 509,259] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
512. 55 <=> v5_algstr_0(sK0) [avatar definition]
514. v5_algstr_0(sK0) <- (55) [avatar component clause 512] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
515. 55 | ~8 | 10 | ~26 [avatar split clause 510,340,257,247,512]
542. ~v4_algstr_1(sK0) | v2_struct_0(sK0) | v3_algstr_1(sK0) <- (26) [resolution 162,342] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
549. v2_struct_0(sK0) | v3_algstr_1(sK0) <- (8, 26) [subsumption resolution 542,249] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
550. v3_algstr_1(sK0) <- (8, ~10, 26) [subsumption resolution 549,259] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
552. 61 <=> v3_algstr_1(sK0) [avatar definition]
554. v3_algstr_1(sK0) <- (61) [avatar component clause 552] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
555. 61 | ~8 | 10 | ~26 [avatar split clause 550,340,257,247,552]
660. ~v4_vectsp_1(sK0) | v2_struct_0(sK0) | v6_vectsp_1(sK0) <- (34) [resolution 184,389] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
673. v2_struct_0(sK0) | v6_vectsp_1(sK0) <- (3, 34) [subsumption resolution 660,224] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
674. v6_vectsp_1(sK0) <- (3, ~10, 34) [subsumption resolution 673,259] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
676. 81 <=> v6_vectsp_1(sK0) [avatar definition]
678. v6_vectsp_1(sK0) <- (81) [avatar component clause 676] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
679. 81 | ~3 | 10 | ~34 [avatar split clause 674,387,257,222,676]
787. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~l6_algstr_0(sK0) | ~v4_vectsp_1(sK0) | ~v2_vectsp_1(sK0) | ~v4_rlvect_1(sK0) | ~v7_algstr_1(sK0) | ~v4_algstr_1(sK0) | ~v1_algstr_1(sK0) | v2_struct_0(sK0) <- (~1) [superposition 214,186] {a:1,w:41,wCS:4100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
794. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~v4_vectsp_1(sK0) | ~v2_vectsp_1(sK0) | ~v4_rlvect_1(sK0) | ~v7_algstr_1(sK0) | ~v4_algstr_1(sK0) | ~v1_algstr_1(sK0) | v2_struct_0(sK0) <- (~1, 2) [subsumption resolution 787,219] {a:1,w:39,wCS:3900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
795. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~v2_vectsp_1(sK0) | ~v4_rlvect_1(sK0) | ~v7_algstr_1(sK0) | ~v4_algstr_1(sK0) | ~v1_algstr_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 3) [subsumption resolution 794,224] {a:1,w:37,wCS:3700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
796. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~v4_rlvect_1(sK0) | ~v7_algstr_1(sK0) | ~v4_algstr_1(sK0) | ~v1_algstr_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4) [subsumption resolution 795,229] {a:1,w:35,wCS:3500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
797. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~v7_algstr_1(sK0) | ~v4_algstr_1(sK0) | ~v1_algstr_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4, 5) [subsumption resolution 796,234] {a:1,w:33,wCS:3300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
798. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~v4_algstr_1(sK0) | ~v1_algstr_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4, 5, 7) [subsumption resolution 797,244] {a:1,w:31,wCS:3100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
799. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~v1_algstr_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4, 5, 7, 8) [subsumption resolution 798,249] {a:1,w:29,wCS:2900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
800. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~1, 2, 3, 4, 5, 7, 8, 9) [subsumption resolution 799,254] {a:1,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
801. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) <- (~1, 2, 3, 4, 5, 7, 8, 9, ~10) [subsumption resolution 800,259] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
803. 91 <=> m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) [avatar definition]
804. m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) <- (91) [avatar component clause 803] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
805. ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) <- (~91) [avatar component clause 803] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
807. 92 <=> m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) [avatar definition]
808. m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) <- (92) [avatar component clause 807] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
809. ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) <- (~92) [avatar component clause 807] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
811. 93 <=> k5_struct_0(sK0) = k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) [avatar definition]
813. k5_struct_0(sK0) != k1_algstr_2(sK0,k6_algstr_0(sK0,k5_struct_0(sK0),k1_algstr_2(sK0,k5_struct_0(sK0)))) <- (~93) [avatar component clause 811] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
814. ~91 | ~92 | ~93 | 1 | ~2 | ~3 | ~4 | ~5 | ~7 | ~8 | ~9 | 10 [avatar split clause 801,257,252,247,242,232,227,222,217,212,811,807,803]
816. ~l3_struct_0(sK0) <- (~91) [resolution 805,171] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
818. $false <- (40, ~91) [subsumption resolution 816,425] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
819. ~40 | 91 [avatar contradiction clause 818]
850. ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~l2_algstr_0(sK0) | ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (~92) [resolution 809,195] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
852. ~l2_algstr_0(sK0) | ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (91, ~92) [subsumption resolution 850,804] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
853. ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (26, 91, ~92) [subsumption resolution 852,342] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
854. ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (26, 61, 91, ~92) [subsumption resolution 853,554] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
855. v2_struct_0(sK0) <- (26, 55, 61, 91, ~92) [subsumption resolution 854,514] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
856. $false <- (~10, 26, 55, 61, 91, ~92) [subsumption resolution 855,259] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
857. 10 | ~26 | ~55 | ~61 | ~91 | 92 [avatar contradiction clause 856]
887. k5_struct_0(sK0) != k1_algstr_2(sK0,k1_algstr_2(sK0,k5_struct_0(sK0))) | ~m1_subset_1(k1_algstr_2(sK0,k5_struct_0(sK0)),u1_struct_0(sK0)) | ~v6_vectsp_1(sK0) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (~93) [superposition 813,189] {a:2,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
888. k5_struct_0(sK0) != k1_algstr_2(sK0,k1_algstr_2(sK0,k5_struct_0(sK0))) | ~v6_vectsp_1(sK0) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (92, ~93) [subsumption resolution 887,808] {a:2,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
889. k5_struct_0(sK0) != k1_algstr_2(sK0,k1_algstr_2(sK0,k5_struct_0(sK0))) | ~l4_algstr_0(sK0) | v2_struct_0(sK0) <- (81, 92, ~93) [subsumption resolution 888,678] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
890. k5_struct_0(sK0) != k1_algstr_2(sK0,k1_algstr_2(sK0,k5_struct_0(sK0))) | v2_struct_0(sK0) <- (34, 81, 92, ~93) [subsumption resolution 889,389] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
891. k5_struct_0(sK0) != k1_algstr_2(sK0,k1_algstr_2(sK0,k5_struct_0(sK0))) <- (~10, 34, 81, 92, ~93) [subsumption resolution 890,259] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
893. 100 <=> k5_struct_0(sK0) = k1_algstr_2(sK0,k1_algstr_2(sK0,k5_struct_0(sK0))) [avatar definition]
895. k5_struct_0(sK0) != k1_algstr_2(sK0,k1_algstr_2(sK0,k5_struct_0(sK0))) <- (~100) [avatar component clause 893] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
896. ~100 | 10 | ~34 | ~81 | ~92 | 93 [avatar split clause 891,811,807,676,387,257,893]
958. k5_struct_0(sK0) != k5_struct_0(sK0) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~l2_algstr_0(sK0) | ~v2_rlvect_1(sK0) | ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (~100) [superposition 895,187] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
959. ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~l2_algstr_0(sK0) | ~v2_rlvect_1(sK0) | ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (~100) [trivial inequality removal 958] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
961. ~l2_algstr_0(sK0) | ~v2_rlvect_1(sK0) | ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (91, ~100) [subsumption resolution 959,804] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
962. ~v2_rlvect_1(sK0) | ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (26, 91, ~100) [subsumption resolution 961,342] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
963. ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (6, 26, 91, ~100) [subsumption resolution 962,239] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
964. ~v5_algstr_0(sK0) | v2_struct_0(sK0) <- (6, 26, 61, 91, ~100) [subsumption resolution 963,554] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
965. v2_struct_0(sK0) <- (6, 26, 55, 61, 91, ~100) [subsumption resolution 964,514] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
966. $false <- (6, ~10, 26, 55, 61, 91, ~100) [subsumption resolution 965,259] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
967. ~6 | 10 | ~26 | ~55 | ~61 | ~91 | 100 [avatar contradiction clause 966]
968. $false [avatar sat refutation 215,220,225,230,235,240,245,250,255,260,343,355,390,408,446,515,555,679,814,819,857,896,967] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:674,thDist:-674}
% SZS output end Proof for algstr_2__t8_algstr_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 49
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 70
% Generated clauses: 388
% Active clauses: 146
% Passive clauses: 173
% Final active clauses: 144
% Final passive clauses: 26
% Main loop iterations started: 147
% 
% >>> Simplifying Inferences
% Duplicate literals: 23
% Trivial inequalities: 1
% Fw subsumption resolutions: 81
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 120
% Unit resulting resolution: 3
% Forward superposition: 11
% Backward superposition: 5
% 
% >>> AVATAR
% Split clauses: 24
% Split components: 67
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 107
% SAT solver unit clauses: 24
% SAT solver binary clauses: 34
% 
% Memory used [KB]: 5628
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% superposition: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

