i: [1,50,0,2,1,1,0,0] t45_xxreal_2
a: [3,73,0,2,1,50]
s: 73
i: [1,49,0,2,1,1,0,0] t45_xxreal_2
a: [3,78,0,2,1,49]
s: 78
i: [1,47,0,2,1,1,0,0] t45_xxreal_2
a: [3,88,0,2,1,47]
s: 88
i: [1,46,0,2,1,1,0,0] t45_xxreal_2
a: [3,93,0,2,1,46]
s: 93
i: [1,48,0,3,1,1,0,0] t45_xxreal_2
a: [3,83,0,3,1,48]
s: 83
i: [1,51,0,3,1,0,0,255] fc1_finsub_1
s: 51
i: [1,69,0,3,1,0,0,1] reflexivity_r1_tarski
s: 69
i: [1,52,0,4,2,0,0,255] cc1_finsub_1
s: 52
d: [2,95,1,3,1,0,40,52,51]
s: 95
i: [1,53,0,4,2,0,0,255] cc1_finsub_1
s: 53
i: [1,68,0,4,1,0,0,255] existence_m1_subset_1
s: 68
d: [2,96,1,3,1,0,40,53,51]
s: 96
i: [1,54,0,6,3,0,0,1] d11_xxreal_2
s: 54
d: [2,97,1,4,2,1,40,54,78]
d: [2,98,1,2,1,2,21,97,88]
a: [3,102,1,2,1,98]
s: 102
i: [1,55,0,6,3,0,0,1] d11_xxreal_2
s: 55
i: [1,60,0,6,3,0,0,255] cc2_finsub_1
s: 60
d: [2,104,1,4,2,1,40,55,78]
d: [2,105,1,2,1,2,21,104,88]
a: [3,109,1,2,1,105]
s: 109
i: [1,64,0,6,3,0,0,1] cc4_xxreal_2
s: 64
i: [1,65,0,6,3,0,0,1] cc4_xxreal_2
s: 65
i: [1,66,0,7,2,0,0,255] t3_subset
s: 66
i: [1,67,0,7,2,0,0,255] t3_subset
s: 67
d: [2,112,1,5,1,0,40,66,68]
s: 112
i: [1,56,0,8,4,0,0,1] d11_xxreal_2
s: 56
i: [1,59,0,8,3,0,0,255] cc20_membered
s: 59
i: [1,62,0,8,4,0,0,1] cc5_xxreal_2
s: 62
i: [1,57,0,11,5,0,0,2] t43_xxreal_2
s: 57
d: [2,117,1,8,4,1,40,57,83]
d: [2,121,1,6,3,2,21,117,109]
d: [2,122,1,4,2,3,21,121,88]
d: [2,123,1,2,1,4,21,122,93]
a: [3,127,1,2,1,123]
s: 127
i: [1,58,0,11,5,0,0,2] t44_xxreal_2
s: 58
d: [2,129,2,6,3,1,40,127,56]
d: [2,130,2,4,2,2,21,129,73]
d: [2,131,2,2,1,3,21,130,93]
d: [2,137,1,8,4,1,40,58,83]
d: [2,141,1,6,3,2,21,137,102]
d: [2,142,1,4,2,3,21,141,88]
d: [2,143,1,2,1,4,21,142,93]
f: 143,131,123,105,98,46,47,48,49,50
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_2__t45_xxreal_2
% SZS output start Proof for xxreal_2__t45_xxreal_2
1. ! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => ((v5_xxreal_2(X1) & r1_tarski(X0,X1)) => v5_xxreal_2(X0)))) [input t45_xxreal_2]
2. ~! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => ((v5_xxreal_2(X1) & r1_tarski(X0,X1)) => v5_xxreal_2(X0)))) [negated conjecture 1]
9. ! [X0] : (v2_membered(X0) => (v5_xxreal_2(X0) <=> (v4_xxreal_2(X0) & v3_xxreal_2(X0)))) [input d11_xxreal_2]
17. ! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => ((v4_xxreal_2(X1) & r1_tarski(X0,X1)) => v4_xxreal_2(X0)))) [input t43_xxreal_2]
18. ! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => ((v3_xxreal_2(X1) & r1_tarski(X0,X1)) => v3_xxreal_2(X0)))) [input t44_xxreal_2]
23. ? [X0] : (? [X1] : ((~v5_xxreal_2(X0) & (v5_xxreal_2(X1) & r1_tarski(X0,X1))) & v2_membered(X1)) & v2_membered(X0)) [ennf transformation 2]
24. ? [X0] : (? [X1] : (~v5_xxreal_2(X0) & v5_xxreal_2(X1) & r1_tarski(X0,X1) & v2_membered(X1)) & v2_membered(X0)) [flattening 23]
26. ! [X0] : ((v5_xxreal_2(X0) <=> (v4_xxreal_2(X0) & v3_xxreal_2(X0))) | ~v2_membered(X0)) [ennf transformation 9]
27. ! [X0] : (! [X1] : ((v4_xxreal_2(X0) | (~v4_xxreal_2(X1) | ~r1_tarski(X0,X1))) | ~v2_membered(X1)) | ~v2_membered(X0)) [ennf transformation 17]
28. ! [X0] : (! [X1] : (v4_xxreal_2(X0) | ~v4_xxreal_2(X1) | ~r1_tarski(X0,X1) | ~v2_membered(X1)) | ~v2_membered(X0)) [flattening 27]
29. ! [X0] : (! [X1] : ((v3_xxreal_2(X0) | (~v3_xxreal_2(X1) | ~r1_tarski(X0,X1))) | ~v2_membered(X1)) | ~v2_membered(X0)) [ennf transformation 18]
30. ! [X0] : (! [X1] : (v3_xxreal_2(X0) | ~v3_xxreal_2(X1) | ~r1_tarski(X0,X1) | ~v2_membered(X1)) | ~v2_membered(X0)) [flattening 29]
38. ? [X0] : (? [X1] : (~v5_xxreal_2(X0) & v5_xxreal_2(X1) & r1_tarski(X0,X1) & v2_membered(X1)) & v2_membered(X0)) => (? [X1] : (~v5_xxreal_2(sK0) & v5_xxreal_2(X1) & r1_tarski(sK0,X1) & v2_membered(X1)) & v2_membered(sK0)) [choice axiom]
39. ? [X1] : (~v5_xxreal_2(sK0) & v5_xxreal_2(X1) & r1_tarski(sK0,X1) & v2_membered(X1)) => (~v5_xxreal_2(sK0) & v5_xxreal_2(sK1) & r1_tarski(sK0,sK1) & v2_membered(sK1)) [choice axiom]
40. (~v5_xxreal_2(sK0) & v5_xxreal_2(sK1) & r1_tarski(sK0,sK1) & v2_membered(sK1)) & v2_membered(sK0) [skolemisation 24,39,38]
41. ! [X0] : (((v5_xxreal_2(X0) | (~v4_xxreal_2(X0) | ~v3_xxreal_2(X0))) & ((v4_xxreal_2(X0) & v3_xxreal_2(X0)) | ~v5_xxreal_2(X0))) | ~v2_membered(X0)) [nnf transformation 26]
42. ! [X0] : (((v5_xxreal_2(X0) | ~v4_xxreal_2(X0) | ~v3_xxreal_2(X0)) & ((v4_xxreal_2(X0) & v3_xxreal_2(X0)) | ~v5_xxreal_2(X0))) | ~v2_membered(X0)) [flattening 41]
46. v2_membered(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
47. v2_membered(sK1) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. r1_tarski(sK0,sK1) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. v5_xxreal_2(sK1) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. ~v5_xxreal_2(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. ~v5_xxreal_2(X0) | v3_xxreal_2(X0) | ~v2_membered(X0) [cnf transformation 42] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. ~v5_xxreal_2(X0) | v4_xxreal_2(X0) | ~v2_membered(X0) [cnf transformation 42] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. ~v4_xxreal_2(X0) | v5_xxreal_2(X0) | ~v3_xxreal_2(X0) | ~v2_membered(X0) [cnf transformation 42] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. ~r1_tarski(X0,X1) | ~v4_xxreal_2(X1) | v4_xxreal_2(X0) | ~v2_membered(X1) | ~v2_membered(X0) [cnf transformation 28] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
58. ~r1_tarski(X0,X1) | ~v3_xxreal_2(X1) | v3_xxreal_2(X0) | ~v2_membered(X1) | ~v2_membered(X0) [cnf transformation 30] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
71. 1 <=> v5_xxreal_2(sK0) [avatar definition]
73. ~v5_xxreal_2(sK0) <- (~1) [avatar component clause 71] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
74. ~1 [avatar split clause 50,71]
76. 2 <=> v5_xxreal_2(sK1) [avatar definition]
78. v5_xxreal_2(sK1) <- (2) [avatar component clause 76] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
79. 2 [avatar split clause 49,76]
81. 3 <=> r1_tarski(sK0,sK1) [avatar definition]
83. r1_tarski(sK0,sK1) <- (3) [avatar component clause 81] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
84. 3 [avatar split clause 48,81]
86. 4 <=> v2_membered(sK1) [avatar definition]
88. v2_membered(sK1) <- (4) [avatar component clause 86] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. 4 [avatar split clause 47,86]
91. 5 <=> v2_membered(sK0) [avatar definition]
93. v2_membered(sK0) <- (5) [avatar component clause 91] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
94. 5 [avatar split clause 46,91]
97. v3_xxreal_2(sK1) | ~v2_membered(sK1) <- (2) [resolution 54,78] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
98. v3_xxreal_2(sK1) <- (2, 4) [subsumption resolution 97,88] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
100. 6 <=> v3_xxreal_2(sK1) [avatar definition]
102. v3_xxreal_2(sK1) <- (6) [avatar component clause 100] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
103. 6 | ~2 | ~4 [avatar split clause 98,86,76,100]
104. v4_xxreal_2(sK1) | ~v2_membered(sK1) <- (2) [resolution 55,78] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
105. v4_xxreal_2(sK1) <- (2, 4) [subsumption resolution 104,88] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
107. 7 <=> v4_xxreal_2(sK1) [avatar definition]
109. v4_xxreal_2(sK1) <- (7) [avatar component clause 107] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
110. 7 | ~2 | ~4 [avatar split clause 105,86,76,107]
117. ~v4_xxreal_2(sK1) | v4_xxreal_2(sK0) | ~v2_membered(sK1) | ~v2_membered(sK0) <- (3) [resolution 57,83] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
121. v4_xxreal_2(sK0) | ~v2_membered(sK1) | ~v2_membered(sK0) <- (3, 7) [subsumption resolution 117,109] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
122. v4_xxreal_2(sK0) | ~v2_membered(sK0) <- (3, 4, 7) [subsumption resolution 121,88] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
123. v4_xxreal_2(sK0) <- (3, 4, 5, 7) [subsumption resolution 122,93] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
125. 8 <=> v4_xxreal_2(sK0) [avatar definition]
127. v4_xxreal_2(sK0) <- (8) [avatar component clause 125] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
128. 8 | ~3 | ~4 | ~5 | ~7 [avatar split clause 123,107,91,86,81,125]
129. v5_xxreal_2(sK0) | ~v3_xxreal_2(sK0) | ~v2_membered(sK0) <- (8) [resolution 127,56] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
130. ~v3_xxreal_2(sK0) | ~v2_membered(sK0) <- (~1, 8) [subsumption resolution 129,73] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
131. ~v3_xxreal_2(sK0) <- (~1, 5, 8) [subsumption resolution 130,93] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
133. 9 <=> v3_xxreal_2(sK0) [avatar definition]
136. ~9 | 1 | ~5 | ~8 [avatar split clause 131,125,91,71,133]
137. ~v3_xxreal_2(sK1) | v3_xxreal_2(sK0) | ~v2_membered(sK1) | ~v2_membered(sK0) <- (3) [resolution 58,83] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
141. v3_xxreal_2(sK0) | ~v2_membered(sK1) | ~v2_membered(sK0) <- (3, 6) [subsumption resolution 137,102] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
142. v3_xxreal_2(sK0) | ~v2_membered(sK0) <- (3, 4, 6) [subsumption resolution 141,88] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
143. v3_xxreal_2(sK0) <- (3, 4, 5, 6) [subsumption resolution 142,93] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
144. 9 | ~3 | ~4 | ~5 | ~6 [avatar split clause 143,100,91,86,81,133]
145. $false [avatar sat refutation 74,79,84,89,94,103,110,128,136,144] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:61,thDist:-61}
% SZS output end Proof for xxreal_2__t45_xxreal_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% >>> Preprocessing
% Introduced skolems: 3
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 62
% Active clauses: 25
% Passive clauses: 33
% Final active clauses: 25
% Final passive clauses: 5
% Main loop iterations started: 28
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 17
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

