i: [1,39,0,2,1,1,0,0] t25_funcop_1
a: [3,61,0,2,1,39]
s: 61
i: [1,38,0,2,1,1,0,0] t25_funcop_1
a: [3,66,0,2,1,38]
s: 66
i: [1,37,0,2,1,1,0,0] t25_funcop_1
a: [3,71,0,2,1,37]
s: 71
i: [1,36,0,2,1,1,0,0] t25_funcop_1
a: [3,76,0,2,1,36]
s: 76
i: [1,35,0,2,1,1,0,0] t25_funcop_1
a: [3,81,0,2,1,35]
s: 81
i: [1,34,0,2,1,1,0,0] t25_funcop_1
a: [3,86,0,2,1,34]
s: 86
i: [1,33,0,2,1,1,0,0] t25_funcop_1
a: [3,91,0,2,1,33]
s: 91
i: [1,32,0,2,1,1,0,0] t25_funcop_1
a: [3,96,0,2,1,32]
s: 96
i: [1,52,0,2,1,0,0,1] rc1_funct_1
a: [3,101,0,2,1,52]
s: 101
i: [1,51,0,2,1,0,0,1] rc1_funct_1
a: [3,106,0,2,1,51]
s: 106
i: [1,40,0,15,1,1,0,0] t25_funcop_1
a: [3,56,0,15,1,40]
s: 56
i: [1,44,0,4,1,0,0,1] dt_k3_relat_1
s: 44
i: [1,45,0,12,5,0,0,2] dt_k13_funct_3
s: 45
i: [1,46,0,12,5,0,0,2] dt_k13_funct_3
s: 46
i: [1,47,0,12,5,0,0,1] fc2_funct_1
s: 47
i: [1,48,0,12,5,0,0,1] fc2_funct_1
s: 48
i: [1,41,0,17,4,0,0,1] t36_relat_1
s: 41
i: [1,49,0,17,7,0,0,1] fc5_funcop_1
s: 49
i: [1,50,0,17,7,0,0,1] fc5_funcop_1
s: 50
i: [1,42,0,22,7,0,0,1] d3_funcop_1
s: 42
i: [1,43,0,25,7,0,0,2] t55_funct_3
s: 43
d: [2,109,1,24,8,0,44,48,41]
d: [2,111,1,22,7,0,18,109]
d: [2,114,1,18,6,0,21,111,44]
s: 114
d: [2,126,1,34,11,0,44,46,43]
d: [2,142,1,30,10,0,21,126,44]
d: [2,143,1,26,9,0,21,142,44]
d: [2,144,1,22,8,0,21,143,48]
d: [2,145,1,18,7,0,21,144,48]
s: 145
d: [2,108,1,29,7,0,44,41,41]
d: [2,112,1,27,6,0,18,108]
d: [2,113,1,23,5,0,21,112,44]
s: 113
d: [2,115,1,34,10,0,44,41,42]
d: [2,119,1,32,9,0,18,115]
d: [2,120,1,28,8,0,21,119,45]
s: 120
d: [2,125,1,48,13,0,44,42,43]
d: [2,138,1,44,12,0,21,125,44]
d: [2,139,1,40,11,0,21,138,44]
d: [2,140,1,36,10,0,21,139,48]
d: [2,141,1,32,9,0,21,140,48]
s: 141
d: [2,272,2,30,9,1,44,56,141]
d: [2,297,2,28,8,2,21,272,61]
d: [2,298,2,26,7,3,21,297,66]
d: [2,299,2,24,6,4,21,298,71]
d: [2,300,2,22,5,5,21,299,76]
d: [2,301,2,20,4,6,21,300,81]
d: [2,302,2,18,3,7,21,301,86]
d: [2,303,2,16,2,8,21,302,91]
d: [2,304,2,14,1,9,21,303,96]
a: [3,308,2,14,1,304]
s: 308
d: [2,122,1,46,13,0,44,43,42]
d: [2,130,1,42,11,0,18,122]
d: [2,133,1,38,10,0,21,130,46]
d: [2,134,1,34,9,0,21,133,45]
s: 134
d: [2,314,3,22,4,1,44,308,41]
d: [2,315,3,20,3,2,21,314,66]
d: [2,316,3,18,2,3,21,315,76]
a: [3,320,3,4,1,316]
s: 320
d: [2,363,4,0,0,5,52,96,91,86,81,320,45]
e: 363
d: [2,364,4,8,4,1,40,320,45]
d: [2,366,4,6,3,2,21,364,81]
d: [2,367,4,4,2,3,21,366,86]
d: [2,368,4,2,1,4,21,367,91]
d: [2,369,4,0,0,5,21,368,96]
e: 369
r: 320
d: [2,124,1,46,13,0,44,43,42]
d: [2,128,1,42,11,0,18,124]
d: [2,136,1,38,10,0,21,128,46]
d: [2,137,1,34,9,0,21,136,45]
s: 137
d: [2,121,1,41,10,0,44,43,41]
d: [2,131,1,39,9,0,18,121]
d: [2,132,1,35,8,0,21,131,44]
s: 132
a: [3,319,3,4,1,316]
s: 319
d: [2,123,1,41,10,0,44,43,41]
d: [2,129,1,39,9,0,18,123]
d: [2,135,1,35,8,0,21,129,44]
s: 135
a: [3,324,3,14,1,316]
s: 324
d: [2,550,4,25,7,1,44,324,42]
d: [2,551,4,12,6,1,19,550]
d: [2,552,4,10,5,2,21,551,81]
d: [2,553,4,8,4,3,21,552,86]
d: [2,554,4,6,3,4,21,553,91]
d: [2,555,4,4,2,5,21,554,96]
d: [2,556,4,2,1,6,21,555,61]
d: [2,557,4,0,0,7,21,556,66]
e: 557
f: 557,369,316,304,32,33,34,35,36,37,38,39,40
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funcop_1__t25_funcop_1
% SZS output start Proof for funcop_1__t25_funcop_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_funct_1(X3) & v1_relat_1(X3)) => k3_relat_1(X2,k3_funcop_1(X3,X0,X1)) = k3_funcop_1(X3,k3_relat_1(X2,X0),k3_relat_1(X2,X1)))))) [input t25_funcop_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_funct_1(X3) & v1_relat_1(X3)) => k3_relat_1(X2,k3_funcop_1(X3,X0,X1)) = k3_funcop_1(X3,k3_relat_1(X2,X0),k3_relat_1(X2,X1)))))) [negated conjecture 1]
3. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => k3_funcop_1(X0,X1,X2) = k3_relat_1(k13_funct_3(X1,X2),X0)))) [input d3_funcop_1]
4. ! [X0,X1] : ((v1_funct_1(X1) & v1_relat_1(X1) & v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k13_funct_3(X0,X1)) & v1_relat_1(k13_funct_3(X0,X1)))) [input dt_k13_funct_3]
6. ! [X0,X1] : v1_relat_1(k3_relat_1(X0,X1)) [input dt_k3_relat_1]
7. ! [X0,X1] : ((v1_funct_1(X1) & v1_relat_1(X1) & v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k3_relat_1(X0,X1)) & v1_relat_1(k3_relat_1(X0,X1)))) [input fc2_funct_1]
10. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => k3_relat_1(k3_relat_1(X0,X1),X2) = k3_relat_1(X0,k3_relat_1(X1,X2))))) [input t36_relat_1]
11. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => k13_funct_3(k3_relat_1(X2,X0),k3_relat_1(X2,X1)) = k3_relat_1(X2,k13_funct_3(X0,X1))))) [input t55_funct_3]
12. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (k3_relat_1(X2,k3_funcop_1(X3,X0,X1)) != k3_funcop_1(X3,k3_relat_1(X2,X0),k3_relat_1(X2,X1)) & (v1_funct_1(X3) & v1_relat_1(X3))) & (v1_funct_1(X2) & v1_relat_1(X2))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
13. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (k3_relat_1(X2,k3_funcop_1(X3,X0,X1)) != k3_funcop_1(X3,k3_relat_1(X2,X0),k3_relat_1(X2,X1)) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 12]
14. ! [X0] : (! [X1] : (! [X2] : (k3_relat_1(k3_relat_1(X0,X1),X2) = k3_relat_1(X0,k3_relat_1(X1,X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 10]
15. ! [X0] : (! [X1] : (! [X2] : (k3_funcop_1(X0,X1,X2) = k3_relat_1(k13_funct_3(X1,X2),X0) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 3]
16. ! [X0] : (! [X1] : (! [X2] : (k3_funcop_1(X0,X1,X2) = k3_relat_1(k13_funct_3(X1,X2),X0) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 15]
17. ! [X0] : (! [X1] : (! [X2] : (k13_funct_3(k3_relat_1(X2,X0),k3_relat_1(X2,X1)) = k3_relat_1(X2,k13_funct_3(X0,X1)) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 11]
18. ! [X0] : (! [X1] : (! [X2] : (k13_funct_3(k3_relat_1(X2,X0),k3_relat_1(X2,X1)) = k3_relat_1(X2,k13_funct_3(X0,X1)) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 17]
19. ! [X0,X1] : ((v1_funct_1(k13_funct_3(X0,X1)) & v1_relat_1(k13_funct_3(X0,X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 4]
20. ! [X0,X1] : ((v1_funct_1(k13_funct_3(X0,X1)) & v1_relat_1(k13_funct_3(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 19]
21. ! [X0,X1] : ((v1_funct_1(k3_relat_1(X0,X1)) & v1_relat_1(k3_relat_1(X0,X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 7]
22. ! [X0,X1] : ((v1_funct_1(k3_relat_1(X0,X1)) & v1_relat_1(k3_relat_1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 21]
25. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (k3_relat_1(X2,k3_funcop_1(X3,X0,X1)) != k3_funcop_1(X3,k3_relat_1(X2,X0),k3_relat_1(X2,X1)) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (? [X3] : (k3_relat_1(X2,k3_funcop_1(X3,sK0,X1)) != k3_funcop_1(X3,k3_relat_1(X2,sK0),k3_relat_1(X2,X1)) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
26. ? [X1] : (? [X2] : (? [X3] : (k3_relat_1(X2,k3_funcop_1(X3,sK0,X1)) != k3_funcop_1(X3,k3_relat_1(X2,sK0),k3_relat_1(X2,X1)) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (k3_relat_1(X2,k3_funcop_1(X3,sK0,sK1)) != k3_funcop_1(X3,k3_relat_1(X2,sK0),k3_relat_1(X2,sK1)) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
27. ? [X2] : (? [X3] : (k3_relat_1(X2,k3_funcop_1(X3,sK0,sK1)) != k3_funcop_1(X3,k3_relat_1(X2,sK0),k3_relat_1(X2,sK1)) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) => (? [X3] : (k3_relat_1(sK2,k3_funcop_1(X3,sK0,sK1)) != k3_funcop_1(X3,k3_relat_1(sK2,sK0),k3_relat_1(sK2,sK1)) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
28. ? [X3] : (k3_relat_1(sK2,k3_funcop_1(X3,sK0,sK1)) != k3_funcop_1(X3,k3_relat_1(sK2,sK0),k3_relat_1(sK2,sK1)) & v1_funct_1(X3) & v1_relat_1(X3)) => (k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_funcop_1(sK3,k3_relat_1(sK2,sK0),k3_relat_1(sK2,sK1)) & v1_funct_1(sK3) & v1_relat_1(sK3)) [choice axiom]
29. (((k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_funcop_1(sK3,k3_relat_1(sK2,sK0),k3_relat_1(sK2,sK1)) & v1_funct_1(sK3) & v1_relat_1(sK3)) & v1_funct_1(sK2) & v1_relat_1(sK2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 13,28,27,26,25]
32. v1_relat_1(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
33. v1_funct_1(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
34. v1_relat_1(sK1) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
35. v1_funct_1(sK1) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
36. v1_relat_1(sK2) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
37. v1_funct_1(sK2) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
38. v1_relat_1(sK3) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
39. v1_funct_1(sK3) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
40. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_funcop_1(sK3,k3_relat_1(sK2,sK0),k3_relat_1(sK2,sK1)) [cnf transformation 29] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
41. k3_relat_1(k3_relat_1(X0,X1),X2) = k3_relat_1(X0,k3_relat_1(X1,X2)) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 14] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. k3_funcop_1(X0,X1,X2) = k3_relat_1(k13_funct_3(X1,X2),X0) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 16] {a:0,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. k13_funct_3(k3_relat_1(X2,X0),k3_relat_1(X2,X1)) = k3_relat_1(X2,k13_funct_3(X0,X1)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:25,wCS:10000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
44. v1_relat_1(k3_relat_1(X0,X1)) [cnf transformation 6] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. v1_relat_1(k13_funct_3(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 20] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
48. v1_funct_1(k3_relat_1(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. 1 <=> k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) = k3_funcop_1(sK3,k3_relat_1(sK2,sK0),k3_relat_1(sK2,sK1)) [avatar definition]
56. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_funcop_1(sK3,k3_relat_1(sK2,sK0),k3_relat_1(sK2,sK1)) <- (~1) [avatar component clause 54] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
57. ~1 [avatar split clause 40,54]
59. 2 <=> v1_funct_1(sK3) [avatar definition]
61. v1_funct_1(sK3) <- (2) [avatar component clause 59] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
62. 2 [avatar split clause 39,59]
64. 3 <=> v1_relat_1(sK3) [avatar definition]
66. v1_relat_1(sK3) <- (3) [avatar component clause 64] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
67. 3 [avatar split clause 38,64]
69. 4 <=> v1_funct_1(sK2) [avatar definition]
71. v1_funct_1(sK2) <- (4) [avatar component clause 69] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
72. 4 [avatar split clause 37,69]
74. 5 <=> v1_relat_1(sK2) [avatar definition]
76. v1_relat_1(sK2) <- (5) [avatar component clause 74] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
77. 5 [avatar split clause 36,74]
79. 6 <=> v1_funct_1(sK1) [avatar definition]
81. v1_funct_1(sK1) <- (6) [avatar component clause 79] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
82. 6 [avatar split clause 35,79]
84. 7 <=> v1_relat_1(sK1) [avatar definition]
86. v1_relat_1(sK1) <- (7) [avatar component clause 84] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
87. 7 [avatar split clause 34,84]
89. 8 <=> v1_funct_1(sK0) [avatar definition]
91. v1_funct_1(sK0) <- (8) [avatar component clause 89] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
92. 8 [avatar split clause 33,89]
94. 9 <=> v1_relat_1(sK0) [avatar definition]
96. v1_relat_1(sK0) <- (9) [avatar component clause 94] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
97. 9 [avatar split clause 32,94]
125. k3_funcop_1(X3,k3_relat_1(X0,X1),k3_relat_1(X0,X2)) = k3_relat_1(k3_relat_1(X0,k13_funct_3(X1,X2)),X3) | ~v1_funct_1(k3_relat_1(X0,X2)) | ~v1_relat_1(k3_relat_1(X0,X2)) | ~v1_funct_1(k3_relat_1(X0,X1)) | ~v1_relat_1(k3_relat_1(X0,X1)) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [superposition 42,43] {a:1,w:48,wCS:19200,sine:1,thAx:0,allAx:2,thDist:-2}
138. k3_funcop_1(X3,k3_relat_1(X0,X1),k3_relat_1(X0,X2)) = k3_relat_1(k3_relat_1(X0,k13_funct_3(X1,X2)),X3) | ~v1_funct_1(k3_relat_1(X0,X2)) | ~v1_funct_1(k3_relat_1(X0,X1)) | ~v1_relat_1(k3_relat_1(X0,X1)) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [subsumption resolution 125,44] {a:1,w:44,wCS:17600,sine:1,thAx:0,allAx:2,thDist:-2}
139. k3_funcop_1(X3,k3_relat_1(X0,X1),k3_relat_1(X0,X2)) = k3_relat_1(k3_relat_1(X0,k13_funct_3(X1,X2)),X3) | ~v1_funct_1(k3_relat_1(X0,X2)) | ~v1_funct_1(k3_relat_1(X0,X1)) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [subsumption resolution 138,44] {a:1,w:40,wCS:16000,sine:1,thAx:0,allAx:2,thDist:-2}
140. k3_funcop_1(X3,k3_relat_1(X0,X1),k3_relat_1(X0,X2)) = k3_relat_1(k3_relat_1(X0,k13_funct_3(X1,X2)),X3) | ~v1_funct_1(k3_relat_1(X0,X1)) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [subsumption resolution 139,48] {a:1,w:36,wCS:14400,sine:1,thAx:0,allAx:2,thDist:-2}
141. k3_funcop_1(X3,k3_relat_1(X0,X1),k3_relat_1(X0,X2)) = k3_relat_1(k3_relat_1(X0,k13_funct_3(X1,X2)),X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [subsumption resolution 140,48] {a:1,w:32,wCS:12800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
272. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [superposition 56,141] {a:2,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
297. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) | ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 272,61] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
298. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3) [subsumption resolution 297,66] {a:2,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
299. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) | ~v1_relat_1(sK2) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 4) [subsumption resolution 298,71] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
300. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 4, 5) [subsumption resolution 299,76] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
301. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 4, 5, 6) [subsumption resolution 300,81] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
302. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 4, 5, 6, 7) [subsumption resolution 301,86] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
303. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) | ~v1_relat_1(sK0) <- (~1, 2, 3, 4, 5, 6, 7, 8) [subsumption resolution 302,91] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
304. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) <- (~1, 2, 3, 4, 5, 6, 7, 8, 9) [subsumption resolution 303,96] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
306. 12 <=> k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) = k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) [avatar definition]
308. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(k3_relat_1(sK2,k13_funct_3(sK0,sK1)),sK3) <- (~12) [avatar component clause 306] {a:2,w:14,wCS:1400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
309. ~12 | 1 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 [avatar split clause 304,94,89,84,79,74,69,64,59,54,306]
314. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(sK2,k3_relat_1(k13_funct_3(sK0,sK1),sK3)) | ~v1_relat_1(sK3) | ~v1_relat_1(k13_funct_3(sK0,sK1)) | ~v1_relat_1(sK2) <- (~12) [superposition 308,41] {a:3,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
315. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(sK2,k3_relat_1(k13_funct_3(sK0,sK1),sK3)) | ~v1_relat_1(k13_funct_3(sK0,sK1)) | ~v1_relat_1(sK2) <- (3, ~12) [subsumption resolution 314,66] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
316. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(sK2,k3_relat_1(k13_funct_3(sK0,sK1),sK3)) | ~v1_relat_1(k13_funct_3(sK0,sK1)) <- (3, 5, ~12) [subsumption resolution 315,76] {a:3,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
318. 13 <=> v1_relat_1(k13_funct_3(sK0,sK1)) [avatar definition]
320. ~v1_relat_1(k13_funct_3(sK0,sK1)) <- (~13) [avatar component clause 318] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
322. 14 <=> k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) = k3_relat_1(sK2,k3_relat_1(k13_funct_3(sK0,sK1),sK3)) [avatar definition]
324. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(sK2,k3_relat_1(k13_funct_3(sK0,sK1),sK3)) <- (~14) [avatar component clause 322] {a:3,w:14,wCS:1400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
325. ~13 | ~14 | ~3 | ~5 | 12 [avatar split clause 316,306,74,64,322,318]
364. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~13) [resolution 320,45] {a:4,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
366. ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (6, ~13) [subsumption resolution 364,81] {a:4,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
367. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (6, 7, ~13) [subsumption resolution 366,86] {a:4,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
368. ~v1_relat_1(sK0) <- (6, 7, 8, ~13) [subsumption resolution 367,91] {a:4,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
369. $false <- (6, 7, 8, 9, ~13) [subsumption resolution 368,96] {a:4,w:0,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
370. ~6 | ~7 | ~8 | ~9 | 13 [avatar contradiction clause 369]
550. k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) != k3_relat_1(sK2,k3_funcop_1(sK3,sK0,sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (~14) [superposition 324,42] {a:4,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
551. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (~14) [trivial inequality removal 550] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
552. ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (6, ~14) [subsumption resolution 551,81] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
553. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (6, 7, ~14) [subsumption resolution 552,86] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
554. ~v1_relat_1(sK0) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (6, 7, 8, ~14) [subsumption resolution 553,91] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
555. ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (6, 7, 8, 9, ~14) [subsumption resolution 554,96] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
556. ~v1_relat_1(sK3) <- (2, 6, 7, 8, 9, ~14) [subsumption resolution 555,61] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
557. $false <- (2, 3, 6, 7, 8, 9, ~14) [subsumption resolution 556,66] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
558. ~2 | ~3 | ~6 | ~7 | ~8 | ~9 | 14 [avatar contradiction clause 557]
559. $false [avatar sat refutation 57,62,67,72,77,82,87,92,97,309,325,370,558] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:137,thDist:-137}
% SZS output end Proof for funcop_1__t25_funcop_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 470
% Active clauses: 33
% Passive clauses: 138
% Final active clauses: 32
% Final passive clauses: 104
% Main loop iterations started: 34
% 
% >>> Simplifying Inferences
% Duplicate literals: 105
% Trivial inequalities: 1
% Fw subsumption resolutions: 243
% 
% >>> Deletion Inferences
% Forward subsumptions: 14
% 
% >>> Generating Inferences
% Binary resolution: 1
% Unit resulting resolution: 1
% Forward superposition: 68
% Backward superposition: 55
% Self superposition: 6
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 11
% 
% Memory used [KB]: 5373
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

