i: [1,74,0,8,2,1,0,0] t14_mcart_1
a: [3,108,0,4,1,74]
s: 108
i: [1,73,0,7,1,1,0,0] t14_mcart_1
a: [3,117,0,7,1,73]
s: 117
i: [1,75,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,122,0,2,1,75]
s: 122
i: [1,76,0,2,1,0,0,255] fc1_xboole_0
a: [3,127,0,2,1,76]
s: 127
i: [1,97,0,2,1,0,0,2] rc2_xboole_0
a: [3,137,0,2,1,97]
s: 137
i: [1,98,0,2,1,0,0,2] rc1_xboole_0
a: [3,142,0,2,1,98]
s: 142
i: [1,100,0,2,1,0,0,2] rc1_relat_1
a: [3,147,0,2,1,100]
s: 147
i: [1,99,0,2,1,0,0,2] rc1_relat_1
a: [3,152,0,2,1,99]
s: 152
i: [1,101,0,2,1,0,0,255] rc2_relat_1
a: [3,157,0,2,1,101]
s: 157
i: [1,78,0,3,1,0,0,1] fc2_xboole_0
s: 78
i: [1,77,0,3,1,0,0,255] d2_xboole_0
a: [3,132,0,3,1,77]
s: 132
r: 122
i: [1,79,0,4,2,0,0,2] cc1_relat_1
s: 79
i: [1,93,0,4,1,0,0,255] existence_m1_subset_1
s: 93
d: [2,160,1,2,1,1,40,79,142]
a: [3,165,1,2,1,160]
s: 165
i: [1,94,0,4,1,0,0,1] fc6_relat_1
s: 94
i: [1,103,0,4,1,0,0,1] d1_tarski
s: 103
d: [2,161,1,2,1,1,40,79,127]
a: [3,170,1,2,1,161]
s: 170
i: [1,80,0,5,2,0,0,255] t6_boole
s: 80
i: [1,90,0,5,2,0,0,1] t7_boole
s: 90
i: [1,84,0,6,2,0,0,255] t1_subset
s: 84
i: [1,95,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 95
d: [2,175,1,6,1,1,40,90,117]
a: [3,180,1,6,1,175]
s: 180
i: [1,89,0,7,3,0,0,1] t8_boole
s: 89
d: [2,182,1,7,1,1,40,84,117]
a: [3,187,1,7,1,182]
s: 187
i: [1,104,0,7,2,0,0,1] d1_tarski
s: 104
i: [1,83,0,8,3,0,0,255] t2_subset
s: 83
d: [2,189,1,7,1,1,40,95,117]
a: [3,194,1,7,1,189]
s: 194
i: [1,96,0,8,3,0,0,1] fc10_subset_1
s: 96
d: [2,183,1,4,1,0,40,84,103]
s: 183
i: [1,91,0,9,2,0,0,1] t10_mcart_1
s: 91
i: [1,92,0,9,2,0,0,1] t10_mcart_1
s: 92
d: [2,200,1,5,1,1,40,91,117]
a: [3,204,1,5,1,200]
s: 204
d: [2,212,2,0,0,2,52,108,204,104]
e: 212
d: [2,213,2,4,1,1,40,204,104]
d: [2,218,2,0,0,2,21,213,108]
e: 218
r: 108
a: [3,107,0,4,1,74]
s: 107
r: 204
a: [3,112,0,4,1,74]
s: 112
i: [1,87,0,14,3,0,0,1] d1_tarski
s: 87
i: [1,88,0,14,3,0,0,1] d1_tarski
s: 88
d: [2,206,1,5,1,1,40,92,117]
a: [3,210,1,5,1,206]
s: 210
d: [2,238,2,0,0,2,52,112,210,104]
e: 238
d: [2,239,2,4,1,1,40,210,104]
d: [2,244,2,0,0,2,21,239,112]
e: 244
f: 244,213,206,200,73,74
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t14_mcart_1
% SZS output start Proof for mcart_1__t14_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))) => (k2_xtuple_0(X0) = X2 & k1_xtuple_0(X0) = X1)) [input t14_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))) => (k2_xtuple_0(X0) = X2 & k1_xtuple_0(X0) = X1)) [negated conjecture 1]
7. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
25. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [input t10_mcart_1]
31. ~! [X0,X1,X2] : (r2_hidden(X0,k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))) => (k2_xtuple_0(X0) = X2 & k1_xtuple_0(X0) = X1)) [flattening 2]
34. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 7]
37. ! [X0,X1,X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [flattening 25]
41. ? [X0,X1,X2] : ((k2_xtuple_0(X0) != X2 | k1_xtuple_0(X0) != X1) & r2_hidden(X0,k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2)))) [ennf transformation 31]
53. ! [X0,X1,X2] : ((r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1)) | ~r2_hidden(X0,k2_zfmisc_1(X1,X2))) [ennf transformation 37]
57. ? [X0,X1,X2] : ((k2_xtuple_0(X0) != X2 | k1_xtuple_0(X0) != X1) & r2_hidden(X0,k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2)))) => ((sK2 != k2_xtuple_0(sK0) | sK1 != k1_xtuple_0(sK0)) & r2_hidden(sK0,k2_zfmisc_1(k1_tarski(sK1),k1_tarski(sK2)))) [choice axiom]
58. (sK2 != k2_xtuple_0(sK0) | sK1 != k1_xtuple_0(sK0)) & r2_hidden(sK0,k2_zfmisc_1(k1_tarski(sK1),k1_tarski(sK2))) [skolemisation 41,57]
59. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 34]
60. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 59]
61. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
62. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 60,61]
73. r2_hidden(sK0,k2_zfmisc_1(k1_tarski(sK1),k1_tarski(sK2))) [cnf transformation 58] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
74. sK2 != k2_xtuple_0(sK0) | sK1 != k1_xtuple_0(sK0) [cnf transformation 58] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
85. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 62] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
91. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k1_xtuple_0(X0),X1) [cnf transformation 53] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
92. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k2_xtuple_0(X0),X2) [cnf transformation 53] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 85] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
106. 1 <=> sK1 = k1_xtuple_0(sK0) [avatar definition]
110. 2 <=> sK2 = k2_xtuple_0(sK0) [avatar definition]
112. sK2 != k2_xtuple_0(sK0) <- (~2) [avatar component clause 110] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
113. ~1 | ~2 [avatar split clause 74,110,106]
115. 3 <=> r2_hidden(sK0,k2_zfmisc_1(k1_tarski(sK1),k1_tarski(sK2))) [avatar definition]
117. r2_hidden(sK0,k2_zfmisc_1(k1_tarski(sK1),k1_tarski(sK2))) <- (3) [avatar component clause 115] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
118. 3 [avatar split clause 73,115]
200. r2_hidden(k1_xtuple_0(sK0),k1_tarski(sK1)) <- (3) [resolution 91,117] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
202. 17 <=> r2_hidden(k1_xtuple_0(sK0),k1_tarski(sK1)) [avatar definition]
204. r2_hidden(k1_xtuple_0(sK0),k1_tarski(sK1)) <- (17) [avatar component clause 202] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
205. 17 | ~3 [avatar split clause 200,115,202]
206. r2_hidden(k2_xtuple_0(sK0),k1_tarski(sK2)) <- (3) [resolution 92,117] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
208. 18 <=> r2_hidden(k2_xtuple_0(sK0),k1_tarski(sK2)) [avatar definition]
210. r2_hidden(k2_xtuple_0(sK0),k1_tarski(sK2)) <- (18) [avatar component clause 208] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
211. 18 | ~3 [avatar split clause 206,115,208]
213. sK1 = k1_xtuple_0(sK0) <- (17) [resolution 204,104] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
220. 1 | ~17 [avatar split clause 213,202,106]
239. sK2 = k2_xtuple_0(sK0) <- (18) [resolution 210,104] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
244. $false <- (~2, 18) [subsumption resolution 239,112] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
245. 2 | ~18 [avatar contradiction clause 244]
246. $false [avatar sat refutation 113,118,205,211,220,245] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:105,thDist:-105}
% SZS output end Proof for mcart_1__t14_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 90
% Active clauses: 37
% Passive clauses: 48
% Final active clauses: 34
% Final passive clauses: 11
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 2
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 28
% Unit resulting resolution: 2
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 9
% SAT solver binary clauses: 15
% 
% Memory used [KB]: 4989
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

