i: [1,40,0,2,1,1,0,0] t115_zfmisc_1
a: [3,76,0,2,1,40]
s: 76
i: [1,42,0,3,1,1,0,0] t115_zfmisc_1
a: [3,62,0,3,1,42]
s: 62
i: [1,41,0,14,2,1,0,0] t115_zfmisc_1
a: [3,71,0,7,1,41]
s: 71
i: [1,43,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,81,0,2,1,43]
s: 81
i: [1,44,0,2,1,0,0,2] fc1_xboole_0
a: [3,86,0,2,1,44]
s: 86
i: [1,55,0,2,1,0,0,1] rc2_xboole_0
a: [3,96,0,2,1,55]
s: 96
i: [1,56,0,2,1,0,0,1] rc1_xboole_0
a: [3,101,0,2,1,56]
s: 101
i: [1,46,0,3,1,0,0,1] t2_xboole_1
s: 46
i: [1,54,0,3,1,0,0,1] reflexivity_r1_tarski
s: 54
i: [1,45,0,3,1,0,0,255] d2_xboole_0
a: [3,91,0,3,1,45]
s: 91
r: 81
i: [1,47,0,5,2,0,0,2] t6_boole
s: 47
i: [1,57,0,5,1,0,0,1] t90_zfmisc_1
s: 57
i: [1,58,0,5,1,0,0,1] t90_zfmisc_1
s: 58
d: [2,105,1,5,2,0,44,46,47]
s: 105
i: [1,51,0,7,3,0,0,2] t8_boole
s: 51
i: [1,48,0,11,3,0,0,1] t90_zfmisc_1
s: 48
d: [2,109,2,2,1,1,40,105,62]
a: [3,113,2,2,1,109]
s: 113
i: [1,52,0,15,3,0,0,1] t114_zfmisc_1
s: 52
d: [2,117,1,8,2,1,40,52,71]
d: [2,124,1,5,1,2,21,117,62]
a: [3,128,1,5,1,124]
s: 128
r: 71
r: 128
a: [3,67,0,7,1,41]
s: 67
a: [3,70,0,7,1,41]
s: 70
d: [2,133,2,9,3,1,44,48,128]
d: [2,134,2,6,2,1,19,133]
a: [3,137,2,3,1,134]
s: 137
i: [1,53,0,15,3,0,0,1] t114_zfmisc_1
s: 53
a: [3,141,2,3,1,134]
s: 141
a: [3,127,1,5,1,124]
s: 127
d: [2,146,1,8,2,1,40,67,52]
a: [3,154,1,5,1,146]
s: 154
r: 67
d: [2,172,3,0,0,3,52,141,137,154,48]
e: 172
d: [2,177,2,9,3,1,44,48,154]
d: [2,178,2,6,2,1,19,177]
d: [2,180,2,3,1,2,21,178,137]
d: [2,181,2,0,0,3,21,180,141]
e: 181
d: [2,162,1,8,2,1,40,53,67]
d: [2,169,1,5,1,2,21,162,62]
f: 181,169,134,124,109,45,43,40,41,42
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t115_zfmisc_1
% SZS output start Proof for zfmisc_1__t115_zfmisc_1
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ! [X2] : ! [X3] : ((r1_tarski(k2_zfmisc_1(X1,X0),k2_zfmisc_1(X3,X2)) | r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3))) => r1_tarski(X1,X3))) [input t115_zfmisc_1]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : ! [X2] : ! [X3] : ((r1_tarski(k2_zfmisc_1(X1,X0),k2_zfmisc_1(X3,X2)) | r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3))) => r1_tarski(X1,X3))) [negated conjecture 1]
4. k1_xboole_0 = o_0_0_xboole_0 [input d2_xboole_0]
7. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
12. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3)) => ((r1_tarski(X1,X3) & r1_tarski(X0,X2)) | k2_zfmisc_1(X0,X1) = k1_xboole_0)) [input t114_zfmisc_1]
13. ! [X0] : r1_tarski(k1_xboole_0,X0) [input t2_xboole_1]
14. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
17. ! [X0] : ! [X1] : (k2_zfmisc_1(X0,X1) = k1_xboole_0 <=> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t90_zfmisc_1]
18. ~! [X0] : (~v1_xboole_0(X0) => ! [X1,X2,X3] : ((r1_tarski(k2_zfmisc_1(X1,X0),k2_zfmisc_1(X3,X2)) | r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3))) => r1_tarski(X1,X3))) [flattening 2]
19. ! [X0,X1] : (k2_zfmisc_1(X0,X1) = k1_xboole_0 <=> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 17]
22. ! [X0,X1,X2,X3] : (r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3)) => ((r1_tarski(X1,X3) & r1_tarski(X0,X2)) | k2_zfmisc_1(X0,X1) = k1_xboole_0)) [flattening 12]
26. ? [X0] : (? [X1,X2,X3] : (~r1_tarski(X1,X3) & (r1_tarski(k2_zfmisc_1(X1,X0),k2_zfmisc_1(X3,X2)) | r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3)))) & ~v1_xboole_0(X0)) [ennf transformation 18]
27. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 14]
29. ! [X0,X1,X2,X3] : (((r1_tarski(X1,X3) & r1_tarski(X0,X2)) | k2_zfmisc_1(X0,X1) = k1_xboole_0) | ~r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3))) [ennf transformation 22]
30. ! [X0,X1,X2,X3] : ((r1_tarski(X1,X3) & r1_tarski(X0,X2)) | k2_zfmisc_1(X0,X1) = k1_xboole_0 | ~r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3))) [flattening 29]
31. ? [X0] : (? [X1,X2,X3] : (~r1_tarski(X1,X3) & (r1_tarski(k2_zfmisc_1(X1,X0),k2_zfmisc_1(X3,X2)) | r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3)))) & ~v1_xboole_0(X0)) => (? [X3,X2,X1] : (~r1_tarski(X1,X3) & (r1_tarski(k2_zfmisc_1(X1,sK0),k2_zfmisc_1(X3,X2)) | r1_tarski(k2_zfmisc_1(sK0,X1),k2_zfmisc_1(X2,X3)))) & ~v1_xboole_0(sK0)) [choice axiom]
32. ? [X3,X2,X1] : (~r1_tarski(X1,X3) & (r1_tarski(k2_zfmisc_1(X1,sK0),k2_zfmisc_1(X3,X2)) | r1_tarski(k2_zfmisc_1(sK0,X1),k2_zfmisc_1(X2,X3)))) => (~r1_tarski(sK1,sK3) & (r1_tarski(k2_zfmisc_1(sK1,sK0),k2_zfmisc_1(sK3,sK2)) | r1_tarski(k2_zfmisc_1(sK0,sK1),k2_zfmisc_1(sK2,sK3)))) [choice axiom]
33. (~r1_tarski(sK1,sK3) & (r1_tarski(k2_zfmisc_1(sK1,sK0),k2_zfmisc_1(sK3,sK2)) | r1_tarski(k2_zfmisc_1(sK0,sK1),k2_zfmisc_1(sK2,sK3)))) & ~v1_xboole_0(sK0) [skolemisation 26,32,31]
34. ! [X0,X1] : ((k2_zfmisc_1(X0,X1) = k1_xboole_0 | (k1_xboole_0 != X1 & k1_xboole_0 != X0)) & ((k1_xboole_0 = X1 | k1_xboole_0 = X0) | k2_zfmisc_1(X0,X1) != k1_xboole_0)) [nnf transformation 19]
35. ! [X0,X1] : ((k2_zfmisc_1(X0,X1) = k1_xboole_0 | (k1_xboole_0 != X1 & k1_xboole_0 != X0)) & (k1_xboole_0 = X1 | k1_xboole_0 = X0 | k2_zfmisc_1(X0,X1) != k1_xboole_0)) [flattening 34]
40. ~v1_xboole_0(sK0) [cnf transformation 33] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. r1_tarski(k2_zfmisc_1(sK1,sK0),k2_zfmisc_1(sK3,sK2)) | r1_tarski(k2_zfmisc_1(sK0,sK1),k2_zfmisc_1(sK2,sK3)) [cnf transformation 33] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. ~r1_tarski(sK1,sK3) [cnf transformation 33] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
43. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 7] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
45. k1_xboole_0 = o_0_0_xboole_0 [cnf transformation 4] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
46. r1_tarski(k1_xboole_0,X0) [cnf transformation 13] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 27] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
48. k2_zfmisc_1(X0,X1) != k1_xboole_0 | k1_xboole_0 = X0 | k1_xboole_0 = X1 [cnf transformation 35] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. ~r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3)) | k2_zfmisc_1(X0,X1) = k1_xboole_0 | r1_tarski(X0,X2) [cnf transformation 30] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. ~r1_tarski(k2_zfmisc_1(X0,X1),k2_zfmisc_1(X2,X3)) | k2_zfmisc_1(X0,X1) = k1_xboole_0 | r1_tarski(X1,X3) [cnf transformation 30] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. 1 <=> r1_tarski(sK1,sK3) [avatar definition]
62. ~r1_tarski(sK1,sK3) <- (~1) [avatar component clause 60] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
63. ~1 [avatar split clause 42,60]
65. 2 <=> r1_tarski(k2_zfmisc_1(sK0,sK1),k2_zfmisc_1(sK2,sK3)) [avatar definition]
67. r1_tarski(k2_zfmisc_1(sK0,sK1),k2_zfmisc_1(sK2,sK3)) <- (2) [avatar component clause 65] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
69. 3 <=> r1_tarski(k2_zfmisc_1(sK1,sK0),k2_zfmisc_1(sK3,sK2)) [avatar definition]
71. r1_tarski(k2_zfmisc_1(sK1,sK0),k2_zfmisc_1(sK3,sK2)) <- (3) [avatar component clause 69] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. 2 | 3 [avatar split clause 41,69,65]
74. 4 <=> v1_xboole_0(sK0) [avatar definition]
77. ~4 [avatar split clause 40,74]
79. 5 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
82. 5 [avatar split clause 43,79]
89. 7 <=> k1_xboole_0 = o_0_0_xboole_0 [avatar definition]
92. 7 [avatar split clause 45,89]
105. r1_tarski(X0,X1) | ~v1_xboole_0(X0) [superposition 46,47] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
109. ~v1_xboole_0(sK1) <- (~1) [resolution 105,62] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
111. 10 <=> v1_xboole_0(sK1) [avatar definition]
114. ~10 | 1 [avatar split clause 109,60,111]
117. k1_xboole_0 = k2_zfmisc_1(sK1,sK0) | r1_tarski(sK1,sK3) <- (3) [resolution 52,71] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
124. k1_xboole_0 = k2_zfmisc_1(sK1,sK0) <- (~1, 3) [subsumption resolution 117,62] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
126. 11 <=> k1_xboole_0 = k2_zfmisc_1(sK1,sK0) [avatar definition]
128. k1_xboole_0 = k2_zfmisc_1(sK1,sK0) <- (11) [avatar component clause 126] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
129. 11 | 1 | ~3 [avatar split clause 124,69,60,126]
133. k1_xboole_0 != k1_xboole_0 | k1_xboole_0 = sK1 | k1_xboole_0 = sK0 <- (11) [superposition 48,128] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
134. k1_xboole_0 = sK1 | k1_xboole_0 = sK0 <- (11) [trivial inequality removal 133] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
136. 12 <=> k1_xboole_0 = sK0 [avatar definition]
137. k1_xboole_0 != sK0 <- (~12) [avatar component clause 136] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
140. 13 <=> k1_xboole_0 = sK1 [avatar definition]
141. k1_xboole_0 != sK1 <- (~13) [avatar component clause 140] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
143. 12 | 13 | ~11 [avatar split clause 134,126,140,136]
144. k1_xboole_0 != sK0 | k1_xboole_0 != o_0_0_xboole_0 | ~v1_xboole_0(o_0_0_xboole_0) | v1_xboole_0(sK0) [theory tautology sat conflict] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:1,thDist:-1}
145. k1_xboole_0 != sK1 | k1_xboole_0 != o_0_0_xboole_0 | ~v1_xboole_0(o_0_0_xboole_0) | v1_xboole_0(sK1) [theory tautology sat conflict] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:1,thDist:-1}
152. 15 <=> k1_xboole_0 = k2_zfmisc_1(sK0,sK1) [avatar definition]
154. k1_xboole_0 = k2_zfmisc_1(sK0,sK1) <- (15) [avatar component clause 152] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
162. k1_xboole_0 = k2_zfmisc_1(sK0,sK1) | r1_tarski(sK1,sK3) <- (2) [resolution 53,67] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
169. k1_xboole_0 = k2_zfmisc_1(sK0,sK1) <- (~1, 2) [subsumption resolution 162,62] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
170. 15 | 1 | ~2 [avatar split clause 169,65,60,152]
177. k1_xboole_0 != k1_xboole_0 | k1_xboole_0 = sK0 | k1_xboole_0 = sK1 <- (15) [superposition 48,154] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
178. k1_xboole_0 = sK0 | k1_xboole_0 = sK1 <- (15) [trivial inequality removal 177] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
180. k1_xboole_0 = sK1 <- (~12, 15) [subsumption resolution 178,137] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
181. $false <- (~12, ~13, 15) [subsumption resolution 180,141] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
182. 12 | 13 | ~15 [avatar contradiction clause 181]
183. $false [avatar sat refutation 63,72,77,82,92,114,129,143,144,145,170,182] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:90,thDist:-90}
% SZS output end Proof for zfmisc_1__t115_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 67
% Active clauses: 26
% Passive clauses: 35
% Final active clauses: 22
% Final passive clauses: 8
% Main loop iterations started: 26
% 
% >>> Simplifying Inferences
% Trivial inequalities: 2
% Fw subsumption resolutions: 4
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 1
% Forward superposition: 6
% Backward superposition: 7
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 7
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

