i: [1,55,0,8,2,1,0,0] l3_zfmisc_1
a: [3,84,0,4,1,55]
s: 84
a: [3,87,0,4,1,55]
s: 87
r: 84
d: [2,121,0,3,1,2,23,84,87]
a: [3,125,0,3,1,121]
s: 125
i: [1,56,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,99,0,2,1,56]
s: 99
i: [1,57,0,2,1,0,0,2] fc1_xboole_0
a: [3,104,0,2,1,57]
s: 104
i: [1,77,0,2,1,0,0,2] rc2_xboole_0
a: [3,114,0,2,1,77]
s: 114
i: [1,78,0,2,1,0,0,2] rc1_xboole_0
a: [3,119,0,2,1,78]
s: 119
i: [1,59,0,3,1,0,0,1] fc2_xboole_0
s: 59
d: [2,127,1,2,1,1,44,59,87]
a: [3,131,1,2,1,127]
s: 131
i: [1,60,0,3,1,0,0,1] t2_xboole_1
s: 60
i: [1,75,0,3,1,0,0,1] reflexivity_r1_tarski
s: 75
d: [2,134,1,0,0,1,40,75,125]
e: 134
d: [2,133,1,0,0,1,52,125,75]
e: 133
r: 87
r: 125
i: [1,54,0,7,2,1,0,0] l3_zfmisc_1
a: [3,93,0,3,1,54]
s: 93
a: [3,83,0,4,1,55]
s: 83
a: [3,88,0,4,1,55]
s: 88
i: [1,79,0,3,1,0,0,1] d10_xboole_0
s: 79
i: [1,80,0,3,1,0,0,1] d10_xboole_0
s: 80
i: [1,58,0,3,1,0,0,255] d2_xboole_0
a: [3,109,0,3,1,58]
s: 109
r: 99
i: [1,61,0,5,1,0,0,2] t4_boole
s: 61
i: [1,62,0,5,1,0,0,2] t3_boole
s: 62
i: [1,63,0,5,2,0,0,2] t6_boole
s: 63
i: [1,73,0,5,2,0,0,2] t7_boole
s: 73
d: [2,142,1,5,2,0,44,60,63]
s: 142
i: [1,64,0,6,2,0,0,1] t3_xboole_1
s: 64
i: [1,76,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 76
i: [1,68,0,7,2,0,0,1] l1_zfmisc_1
s: 68
i: [1,69,0,7,2,0,0,1] l1_zfmisc_1
s: 69
d: [2,148,1,4,1,0,40,68,75]
s: 148
i: [1,72,0,7,3,0,0,2] t8_boole
s: 72
d: [2,152,2,4,1,0,40,148,76]
s: 152
i: [1,70,0,8,2,0,0,2] t37_xboole_1
s: 70
i: [1,71,0,8,2,0,0,2] t37_xboole_1
s: 71
d: [2,155,1,6,2,0,44,70,62]
s: 155
i: [1,67,0,9,3,0,0,1] d10_xboole_0
s: 67
d: [2,164,1,8,2,1,40,67,83]
d: [2,167,1,4,1,2,21,164,88]
a: [3,171,1,4,1,167]
s: 171
i: [1,74,0,12,3,0,0,1] l2_zfmisc_1
s: 74
d: [2,139,1,7,3,0,44,63,63]
s: 139
d: [2,173,2,3,1,1,40,171,69]
a: [3,178,2,3,1,173]
s: 178
d: [2,140,1,7,2,0,44,62,63]
s: 140
d: [2,141,1,7,2,0,44,61,63]
s: 141
d: [2,151,1,7,2,0,40,69,64]
s: 151
d: [2,147,1,8,3,0,44,64,63]
s: 147
d: [2,198,2,6,2,1,44,171,151]
d: [2,203,2,3,1,1,21,198,60]
a: [3,207,2,3,1,203]
s: 207
d: [2,166,1,11,3,0,40,67,69]
s: 166
d: [2,193,2,9,3,0,44,68,151]
d: [2,201,2,6,2,0,21,193,60]
s: 201
d: [2,183,1,13,4,0,44,74,71]
s: 183
d: [2,181,1,14,3,0,40,74,68]
s: 181
d: [2,228,2,9,3,1,40,183,83]
s: 228
d: [2,180,1,18,4,0,40,74,67]
s: 180
d: [2,200,2,6,2,1,44,83,151]
a: [3,212,2,3,1,200]
s: 212
d: [2,257,3,0,0,2,52,93,212,64]
e: 257
d: [2,258,3,0,0,2,52,60,93,212,67]
e: 258
d: [2,259,3,0,0,2,52,60,93,212,67]
e: 259
d: [2,260,3,0,0,3,52,104,93,212,147]
e: 260
d: [2,261,3,3,1,1,40,212,64]
d: [2,269,3,0,0,2,21,261,93]
e: 269
d: [2,262,3,5,2,1,40,212,147]
d: [2,271,3,2,1,2,21,262,93]
d: [2,272,3,0,0,3,21,271,104]
e: 272
d: [2,263,3,6,2,1,40,212,67]
d: [2,274,3,3,1,2,21,263,93]
d: [2,275,3,0,0,2,21,274,60]
e: 275
r: 228
r: 83
r: 93
r: 131
r: 171
r: 178
r: 212
a: [3,92,0,3,1,54]
s: 92
r: 88
s: 84
d: [2,283,0,4,1,2,22,84,92]
d: [2,284,0,0,0,2,21,283,60]
e: 284
i: [1,53,0,11,3,1,0,0] l3_zfmisc_1
d: [2,243,3,6,2,1,40,228,75]
d: [2,246,3,3,1,2,21,243,178]
f: 284,275,246,173,167,133,121,53,54,55
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__l3_zfmisc_1
% SZS output start Proof for zfmisc_1__l3_zfmisc_1
1. ! [X0] : ! [X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [input l3_zfmisc_1]
2. ~! [X0] : ! [X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [negated conjecture 1]
4. ! [X0] : ! [X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [input d10_xboole_0]
12. ! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [input l1_zfmisc_1]
13. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,X1) => (r1_tarski(X0,k4_xboole_0(X1,k1_tarski(X2))) | r2_hidden(X2,X0))) [input l2_zfmisc_1]
16. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
17. ! [X0] : r1_tarski(k1_xboole_0,X0) [input t2_xboole_1]
18. ! [X0] : ! [X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [input t37_xboole_1]
25. ~! [X0,X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [flattening 2]
26. ! [X0,X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [flattening 4]
27. ! [X0,X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [flattening 12]
28. ! [X0,X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [flattening 18]
31. ! [X0,X1,X2] : (r1_tarski(X0,X1) => (r1_tarski(X0,k4_xboole_0(X1,k1_tarski(X2))) | r2_hidden(X2,X0))) [flattening 13]
32. ! [X0] : r1_tarski(X0,X0) [rectify 16]
33. ? [X0,X1] : (r1_tarski(X0,k1_tarski(X1)) <~> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [ennf transformation 25]
38. ! [X0,X1,X2] : ((r1_tarski(X0,k4_xboole_0(X1,k1_tarski(X2))) | r2_hidden(X2,X0)) | ~r1_tarski(X0,X1)) [ennf transformation 31]
39. ! [X0,X1,X2] : (r1_tarski(X0,k4_xboole_0(X1,k1_tarski(X2))) | r2_hidden(X2,X0) | ~r1_tarski(X0,X1)) [flattening 38]
41. ? [X0,X1] : (((k1_tarski(X1) != X0 & k1_xboole_0 != X0) | ~r1_tarski(X0,k1_tarski(X1))) & ((k1_tarski(X1) = X0 | k1_xboole_0 = X0) | r1_tarski(X0,k1_tarski(X1)))) [nnf transformation 33]
42. ? [X0,X1] : (((k1_tarski(X1) != X0 & k1_xboole_0 != X0) | ~r1_tarski(X0,k1_tarski(X1))) & (k1_tarski(X1) = X0 | k1_xboole_0 = X0 | r1_tarski(X0,k1_tarski(X1)))) [flattening 41]
43. ? [X0,X1] : (((k1_tarski(X1) != X0 & k1_xboole_0 != X0) | ~r1_tarski(X0,k1_tarski(X1))) & (k1_tarski(X1) = X0 | k1_xboole_0 = X0 | r1_tarski(X0,k1_tarski(X1)))) => (((sK0 != k1_tarski(sK1) & k1_xboole_0 != sK0) | ~r1_tarski(sK0,k1_tarski(sK1))) & (sK0 = k1_tarski(sK1) | k1_xboole_0 = sK0 | r1_tarski(sK0,k1_tarski(sK1)))) [choice axiom]
44. ((sK0 != k1_tarski(sK1) & k1_xboole_0 != sK0) | ~r1_tarski(sK0,k1_tarski(sK1))) & (sK0 = k1_tarski(sK1) | k1_xboole_0 = sK0 | r1_tarski(sK0,k1_tarski(sK1))) [skolemisation 42,43]
45. ! [X0,X1] : ((X0 = X1 | (~r1_tarski(X1,X0) | ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [nnf transformation 26]
46. ! [X0,X1] : ((X0 = X1 | ~r1_tarski(X1,X0) | ~r1_tarski(X0,X1)) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [flattening 45]
47. ! [X0,X1] : ((r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X1) | ~r1_tarski(k1_tarski(X0),X1))) [nnf transformation 27]
48. ! [X0,X1] : ((k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | k1_xboole_0 != k4_xboole_0(X0,X1))) [nnf transformation 28]
53. sK0 = k1_tarski(sK1) | k1_xboole_0 = sK0 | r1_tarski(sK0,k1_tarski(sK1)) [cnf transformation 44] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. k1_xboole_0 != sK0 | ~r1_tarski(sK0,k1_tarski(sK1)) [cnf transformation 44] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. sK0 != k1_tarski(sK1) | ~r1_tarski(sK0,k1_tarski(sK1)) [cnf transformation 44] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
60. r1_tarski(k1_xboole_0,X0) [cnf transformation 17] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. ~r1_tarski(X1,X0) | X0 = X1 | ~r1_tarski(X0,X1) [cnf transformation 46] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 47] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1) [cnf transformation 48] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
74. r1_tarski(X0,k4_xboole_0(X1,k1_tarski(X2))) | r2_hidden(X2,X0) | ~r1_tarski(X0,X1) [cnf transformation 39] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
75. r1_tarski(X0,X0) [cnf transformation 32] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
82. 1 <=> r1_tarski(sK0,k1_tarski(sK1)) [avatar definition]
83. r1_tarski(sK0,k1_tarski(sK1)) <- (1) [avatar component clause 82] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. ~r1_tarski(sK0,k1_tarski(sK1)) <- (~1) [avatar component clause 82] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
86. 2 <=> sK0 = k1_tarski(sK1) [avatar definition]
87. sK0 = k1_tarski(sK1) <- (2) [avatar component clause 86] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
88. sK0 != k1_tarski(sK1) <- (~2) [avatar component clause 86] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
89. ~1 | ~2 [avatar split clause 55,86,82]
91. 3 <=> k1_xboole_0 = sK0 [avatar definition]
92. k1_xboole_0 = sK0 <- (3) [avatar component clause 91] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
93. k1_xboole_0 != sK0 <- (~3) [avatar component clause 91] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
94. ~1 | ~3 [avatar split clause 54,91,82]
95. 1 | 3 | 2 [avatar split clause 53,86,91,82]
121. ~r1_tarski(sK0,sK0) <- (~1, 2) [backward demodulation 84,87] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
123. 9 <=> r1_tarski(sK0,sK0) [avatar definition]
125. ~r1_tarski(sK0,sK0) <- (~9) [avatar component clause 123] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
126. ~9 | 1 | ~2 [avatar split clause 121,86,82,123]
133. $false <- (~9) [unit resulting resolution 125,75] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
136. 9 [avatar contradiction clause 133]
164. sK0 = k1_tarski(sK1) | ~r1_tarski(k1_tarski(sK1),sK0) <- (1) [resolution 67,83] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
167. ~r1_tarski(k1_tarski(sK1),sK0) <- (1, ~2) [subsumption resolution 164,88] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
169. 11 <=> r1_tarski(k1_tarski(sK1),sK0) [avatar definition]
171. ~r1_tarski(k1_tarski(sK1),sK0) <- (~11) [avatar component clause 169] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
172. ~11 | ~1 | 2 [avatar split clause 167,86,82,169]
173. ~r2_hidden(sK1,sK0) <- (~11) [resolution 171,69] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
176. 12 <=> r2_hidden(sK1,sK0) [avatar definition]
178. ~r2_hidden(sK1,sK0) <- (~12) [avatar component clause 176] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
179. ~12 | 11 [avatar split clause 173,169,176]
183. ~r1_tarski(X2,k1_tarski(X3)) | r2_hidden(X3,X4) | ~r1_tarski(X4,X2) | r1_tarski(X4,k1_xboole_0) [superposition 74,71] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
210. 14 <=> r1_tarski(sK0,k1_xboole_0) [avatar definition]
212. r1_tarski(sK0,k1_xboole_0) <- (14) [avatar component clause 210] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
228. ~r1_tarski(X0,sK0) | r2_hidden(sK1,X0) | r1_tarski(X0,k1_xboole_0) <- (1) [resolution 183,83] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
243. r2_hidden(sK1,sK0) | r1_tarski(sK0,k1_xboole_0) <- (1) [resolution 228,75] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
246. r1_tarski(sK0,k1_xboole_0) <- (1, ~12) [subsumption resolution 243,178] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
247. 14 | ~1 | 12 [avatar split clause 246,176,82,210]
263. k1_xboole_0 = sK0 | ~r1_tarski(k1_xboole_0,sK0) <- (14) [resolution 212,67] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
274. ~r1_tarski(k1_xboole_0,sK0) <- (~3, 14) [subsumption resolution 263,93] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
275. $false <- (~3, 14) [subsumption resolution 274,60] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
276. 3 | ~14 [avatar contradiction clause 275]
283. ~r1_tarski(k1_xboole_0,k1_tarski(sK1)) <- (~1, 3) [forward demodulation 84,92] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
284. $false <- (~1, 3) [subsumption resolution 283,60] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
285. 1 | ~3 [avatar contradiction clause 284]
286. $false [avatar sat refutation 89,94,95,126,136,172,179,247,276,285] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:131,thDist:-131}
% SZS output end Proof for zfmisc_1__l3_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 166
% Active clauses: 47
% Passive clauses: 71
% Final active clauses: 35
% Final passive clauses: 19
% Main loop iterations started: 51
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 11
% Fw demodulations: 2
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Equational tautologies: 6
% Forward subsumptions: 43
% 
% >>> Generating Inferences
% Binary resolution: 48
% Unit resulting resolution: 5
% Forward superposition: 25
% Backward superposition: 21
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 9
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 29
% SAT solver unit clauses: 7
% SAT solver binary clauses: 11
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% superposition: 0.001 s ( own 0.0000 s ) 
% other: 0.002 s ( own 0.001 s ) 

