i: [1,19,0,3,1,1,0,0] l4_jordan21
a: [3,35,0,3,1,19]
s: 35
i: [1,18,0,3,1,1,0,0] l4_jordan21
a: [3,40,0,3,1,18]
s: 40
i: [1,17,0,3,1,1,0,0] l4_jordan21
a: [3,45,0,3,1,17]
s: 45
i: [1,20,0,7,1,1,0,0] l4_jordan21
a: [3,30,0,7,1,20]
s: 30
i: [1,24,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 24
i: [1,26,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 26
d: [2,47,1,3,1,1,40,26,35]
a: [3,53,1,3,1,47]
s: 53
i: [1,25,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 25
r: 30
d: [2,48,1,3,1,1,40,26,40]
a: [3,58,1,3,1,48]
s: 58
d: [2,66,0,7,1,1,23,30,25]
a: [3,70,0,7,1,66]
s: 70
i: [1,22,0,8,2,0,0,1] t70_xboole_1
s: 22
d: [2,49,1,3,1,1,40,26,45]
a: [3,63,1,3,1,49]
s: 63
i: [1,23,0,8,2,0,0,1] t70_xboole_1
s: 23
d: [2,74,1,8,2,0,44,22,25]
s: 74
i: [1,21,0,11,3,0,0,1] t70_xboole_1
s: 21
d: [2,75,1,8,2,0,44,22,25]
s: 75
d: [2,82,1,11,3,0,40,21,26]
s: 82
d: [2,87,2,8,2,1,40,82,70]
d: [2,95,2,5,1,2,21,87,53]
a: [3,99,2,5,1,95]
s: 99
d: [2,101,3,0,0,3,52,58,63,99,21]
e: 101
d: [2,102,3,6,2,1,40,99,21]
d: [2,104,3,3,1,2,21,102,63]
d: [2,105,3,0,0,3,21,104,58]
e: 105
f: 105,95,66,49,48,47,17,18,19,20
% Refutation found. Thanks to Tanya!
% SZS status Theorem for jordan21__l4_jordan21
% SZS output start Proof for jordan21__l4_jordan21
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_xboole_0(X2,X3) & r1_xboole_0(X1,X3) & r1_xboole_0(X0,X3)) => r1_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3)) [input l4_jordan21]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_xboole_0(X2,X3) & r1_xboole_0(X1,X3) & r1_xboole_0(X0,X3)) => r1_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3)) [negated conjecture 1]
3. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
6. ! [X0,X1] : (r1_xboole_0(X0,X1) => r1_xboole_0(X1,X0)) [input symmetry_r1_xboole_0]
7. ! [X0] : ! [X1] : ! [X2] : (~(r1_xboole_0(X0,k2_xboole_0(X1,X2)) & ~(r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1))) & ~(r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1) & ~r1_xboole_0(X0,k2_xboole_0(X1,X2)))) [input t70_xboole_1]
8. ~! [X0,X1,X2,X3] : ((r1_xboole_0(X2,X3) & r1_xboole_0(X1,X3) & r1_xboole_0(X0,X3)) => r1_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3)) [flattening 2]
9. ! [X0,X1,X2] : (~(r1_xboole_0(X0,k2_xboole_0(X1,X2)) & ~(r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1))) & ~(r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1) & ~r1_xboole_0(X0,k2_xboole_0(X1,X2)))) [flattening 7]
11. ? [X0,X1,X2,X3] : (~r1_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3) & (r1_xboole_0(X2,X3) & r1_xboole_0(X1,X3) & r1_xboole_0(X0,X3))) [ennf transformation 8]
12. ? [X0,X1,X2,X3] : (~r1_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3) & r1_xboole_0(X2,X3) & r1_xboole_0(X1,X3) & r1_xboole_0(X0,X3)) [flattening 11]
13. ! [X0,X1,X2] : ((~r1_xboole_0(X0,k2_xboole_0(X1,X2)) | (r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1))) & (~r1_xboole_0(X0,X2) | ~r1_xboole_0(X0,X1) | r1_xboole_0(X0,k2_xboole_0(X1,X2)))) [ennf transformation 9]
14. ! [X0,X1] : (r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 6]
15. ? [X0,X1,X2,X3] : (~r1_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3) & r1_xboole_0(X2,X3) & r1_xboole_0(X1,X3) & r1_xboole_0(X0,X3)) => (~r1_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) & r1_xboole_0(sK2,sK3) & r1_xboole_0(sK1,sK3) & r1_xboole_0(sK0,sK3)) [choice axiom]
16. ~r1_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) & r1_xboole_0(sK2,sK3) & r1_xboole_0(sK1,sK3) & r1_xboole_0(sK0,sK3) [skolemisation 12,15]
17. r1_xboole_0(sK0,sK3) [cnf transformation 16] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
18. r1_xboole_0(sK1,sK3) [cnf transformation 16] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
19. r1_xboole_0(sK2,sK3) [cnf transformation 16] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
20. ~r1_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) [cnf transformation 16] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
21. r1_xboole_0(X0,k2_xboole_0(X1,X2)) | ~r1_xboole_0(X0,X1) | ~r1_xboole_0(X0,X2) [cnf transformation 13] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
25. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
26. ~r1_xboole_0(X0,X1) | r1_xboole_0(X1,X0) [cnf transformation 14] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
28. 1 <=> r1_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) [avatar definition]
30. ~r1_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3) <- (~1) [avatar component clause 28] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
31. ~1 [avatar split clause 20,28]
33. 2 <=> r1_xboole_0(sK2,sK3) [avatar definition]
35. r1_xboole_0(sK2,sK3) <- (2) [avatar component clause 33] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
36. 2 [avatar split clause 19,33]
38. 3 <=> r1_xboole_0(sK1,sK3) [avatar definition]
40. r1_xboole_0(sK1,sK3) <- (3) [avatar component clause 38] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
41. 3 [avatar split clause 18,38]
43. 4 <=> r1_xboole_0(sK0,sK3) [avatar definition]
45. r1_xboole_0(sK0,sK3) <- (4) [avatar component clause 43] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
46. 4 [avatar split clause 17,43]
47. r1_xboole_0(sK3,sK2) <- (2) [resolution 26,35] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
48. r1_xboole_0(sK3,sK1) <- (3) [resolution 26,40] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
49. r1_xboole_0(sK3,sK0) <- (4) [resolution 26,45] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
51. 5 <=> r1_xboole_0(sK3,sK2) [avatar definition]
53. r1_xboole_0(sK3,sK2) <- (5) [avatar component clause 51] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
54. 5 | ~2 [avatar split clause 47,33,51]
56. 6 <=> r1_xboole_0(sK3,sK1) [avatar definition]
58. r1_xboole_0(sK3,sK1) <- (6) [avatar component clause 56] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
59. 6 | ~3 [avatar split clause 48,38,56]
61. 7 <=> r1_xboole_0(sK3,sK0) [avatar definition]
63. r1_xboole_0(sK3,sK0) <- (7) [avatar component clause 61] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
64. 7 | ~4 [avatar split clause 49,43,61]
66. ~r1_xboole_0(k2_xboole_0(sK2,k2_xboole_0(sK0,sK1)),sK3) <- (~1) [backward demodulation 30,25] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. 8 <=> r1_xboole_0(k2_xboole_0(sK2,k2_xboole_0(sK0,sK1)),sK3) [avatar definition]
70. ~r1_xboole_0(k2_xboole_0(sK2,k2_xboole_0(sK0,sK1)),sK3) <- (~8) [avatar component clause 68] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. ~8 | 1 [avatar split clause 66,28,68]
82. r1_xboole_0(k2_xboole_0(X7,X8),X6) | ~r1_xboole_0(X6,X8) | ~r1_xboole_0(X6,X7) [resolution 21,26] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
87. ~r1_xboole_0(sK3,k2_xboole_0(sK0,sK1)) | ~r1_xboole_0(sK3,sK2) <- (~8) [resolution 82,70] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
95. ~r1_xboole_0(sK3,k2_xboole_0(sK0,sK1)) <- (5, ~8) [subsumption resolution 87,53] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
97. 9 <=> r1_xboole_0(sK3,k2_xboole_0(sK0,sK1)) [avatar definition]
99. ~r1_xboole_0(sK3,k2_xboole_0(sK0,sK1)) <- (~9) [avatar component clause 97] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
100. ~9 | ~5 | 8 [avatar split clause 95,68,51,97]
102. ~r1_xboole_0(sK3,sK0) | ~r1_xboole_0(sK3,sK1) <- (~9) [resolution 99,21] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
104. ~r1_xboole_0(sK3,sK1) <- (7, ~9) [subsumption resolution 102,63] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
105. $false <- (6, 7, ~9) [subsumption resolution 104,58] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
106. ~6 | ~7 | 9 [avatar contradiction clause 105]
107. $false [avatar sat refutation 31,36,41,46,54,59,64,71,100,106] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:54,thDist:-54}
% SZS output end Proof for jordan21__l4_jordan21
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 6
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 10
% Generated clauses: 52
% Active clauses: 16
% Passive clauses: 20
% Final active clauses: 15
% Final passive clauses: 2
% Main loop iterations started: 18
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 14
% Unit resulting resolution: 1
% Forward superposition: 12
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 4
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

