i: [1,46,0,5,1,1,0,0] t82_zfmisc_1
a: [3,70,0,5,1,46]
s: 70
i: [1,47,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,75,0,2,1,47]
s: 75
i: [1,48,0,2,1,0,0,255] fc1_xboole_0
a: [3,80,0,2,1,48]
s: 80
i: [1,63,0,2,1,0,0,3] rc2_xboole_0
a: [3,90,0,2,1,63]
s: 90
i: [1,64,0,2,1,0,0,3] rc1_xboole_0
a: [3,95,0,2,1,64]
s: 95
i: [1,61,0,3,1,0,0,1] reflexivity_r1_tarski
s: 61
i: [1,49,0,3,1,0,0,255] d2_xboole_0
a: [3,85,0,3,1,49]
s: 85
r: 75
i: [1,50,0,5,2,0,0,255] t6_boole
s: 50
i: [1,60,0,5,2,0,0,2] t7_boole
s: 60
i: [1,62,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 62
i: [1,51,0,7,2,0,0,1] l42_zfmisc_1
s: 51
i: [1,59,0,7,3,0,0,2] t8_boole
s: 59
i: [1,65,0,7,2,0,0,1] d1_zfmisc_1
s: 65
d: [2,101,1,6,2,0,40,65,60]
s: 101
i: [1,66,0,7,2,0,0,1] d1_zfmisc_1
s: 66
i: [1,53,0,8,2,0,0,1] d3_tarski
s: 53
d: [2,104,1,5,2,0,40,53,60]
s: 104
i: [1,54,0,8,2,0,0,1] d3_tarski
s: 54
d: [2,106,2,2,1,1,40,104,70]
a: [3,110,2,2,1,106]
s: 110
i: [1,52,0,9,3,0,0,1] d3_tarski
s: 52
i: [1,57,0,14,3,0,0,1] d1_zfmisc_1
s: 57
d: [2,98,1,7,3,0,44,50,50]
s: 98
i: [1,58,0,14,3,0,0,1] d1_zfmisc_1
s: 58
d: [2,100,1,7,2,0,40,65,62]
s: 100
d: [2,103,1,8,2,0,40,53,62]
s: 103
d: [2,105,1,10,2,0,40,53,66]
s: 105
d: [2,124,2,8,2,0,40,100,65]
s: 124
d: [2,113,1,10,2,0,40,54,65]
s: 113
d: [2,125,3,5,1,0,40,124,61]
s: 125
d: [2,115,1,10,3,0,40,52,65]
s: 115
d: [2,116,1,11,3,0,40,52,53]
s: 116
d: [2,133,4,4,1,0,40,125,104]
s: 133
d: [2,119,1,11,3,0,40,57,60]
s: 119
d: [2,132,4,6,1,0,40,125,51]
s: 132
d: [2,118,1,14,3,0,40,57,62]
s: 118
d: [2,117,1,17,4,0,40,57,52]
s: 117
d: [2,144,5,7,1,0,40,132,65]
s: 144
d: [2,120,1,17,3,0,40,57,66]
s: 120
d: [2,140,2,8,3,0,40,116,60]
s: 140
d: [2,122,1,17,3,0,40,58,65]
s: 122
d: [2,130,2,9,2,0,40,113,104]
s: 130
d: [2,154,6,8,1,0,40,144,51]
s: 154
d: [2,166,2,16,3,0,40,122,104]
d: [2,169,2,10,2,0,21,166,104]
s: 169
d: [2,126,3,9,2,0,40,124,51]
s: 126
d: [2,135,2,11,3,0,40,115,51]
s: 135
d: [2,139,2,11,3,0,40,116,62]
s: 139
d: [2,173,4,8,2,0,40,126,104]
s: 173
d: [2,129,2,12,2,0,40,113,51]
s: 129
d: [2,179,3,10,2,0,40,129,53]
d: [2,183,3,5,1,0,18,179]
s: 183
d: [2,189,4,0,0,1,40,183,70]
e: 189
d: [2,184,4,0,0,1,52,70,183]
e: 184
f: 184,46
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t82_zfmisc_1
% SZS output start Proof for zfmisc_1__t82_zfmisc_1
1. ! [X0] : r1_tarski(X0,k1_zfmisc_1(k3_tarski(X0))) [input t82_zfmisc_1]
2. ~! [X0] : r1_tarski(X0,k1_zfmisc_1(k3_tarski(X0))) [negated conjecture 1]
4. ! [X0] : ! [X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [input d1_zfmisc_1]
6. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
12. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => r1_tarski(X0,k3_tarski(X1))) [input l42_zfmisc_1]
19. ! [X0,X1] : (r2_hidden(X0,X1) => r1_tarski(X0,k3_tarski(X1))) [flattening 12]
20. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 6]
21. ! [X0,X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [flattening 4]
25. ? [X0] : ~r1_tarski(X0,k1_zfmisc_1(k3_tarski(X0))) [ennf transformation 2]
27. ! [X0,X1] : (r1_tarski(X0,k3_tarski(X1)) | ~r2_hidden(X0,X1)) [ennf transformation 19]
28. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 20]
32. ? [X0] : ~r1_tarski(X0,k1_zfmisc_1(k3_tarski(X0))) => ~r1_tarski(sK0,k1_zfmisc_1(k3_tarski(sK0))) [choice axiom]
33. ~r1_tarski(sK0,k1_zfmisc_1(k3_tarski(sK0))) [skolemisation 25,32]
34. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 28]
35. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 34]
36. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK1(X0,X1),X1) & r2_hidden(sK1(X0,X1),X0))) [choice axiom]
37. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK1(X0,X1),X1) & r2_hidden(sK1(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 35,36]
38. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ~r1_tarski(X2,X0)) & (r1_tarski(X2,X0) | ~r2_hidden(X2,X1))) | k1_zfmisc_1(X0) != X1)) [nnf transformation 21]
39. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [rectify 38]
40. ! [X1,X0] : (? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1))) => ((~r1_tarski(sK2(X0,X1),X0) | ~r2_hidden(sK2(X0,X1),X1)) & (r1_tarski(sK2(X0,X1),X0) | r2_hidden(sK2(X0,X1),X1)))) [choice axiom]
41. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ((~r1_tarski(sK2(X0,X1),X0) | ~r2_hidden(sK2(X0,X1),X1)) & (r1_tarski(sK2(X0,X1),X0) | r2_hidden(sK2(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [skolemisation 39,40]
46. ~r1_tarski(sK0,k1_zfmisc_1(k3_tarski(sK0))) [cnf transformation 33] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. r1_tarski(X0,k3_tarski(X1)) | ~r2_hidden(X0,X1) [cnf transformation 27] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. r2_hidden(sK1(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 37] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
54. ~r2_hidden(sK1(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 37] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
56. r2_hidden(X3,X1) | ~r1_tarski(X3,X0) | k1_zfmisc_1(X0) != X1 [cnf transformation 41] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
65. r2_hidden(X3,k1_zfmisc_1(X0)) | ~r1_tarski(X3,X0) [equality resolution 56] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
68. 1 <=> r1_tarski(sK0,k1_zfmisc_1(k3_tarski(sK0))) [avatar definition]
70. ~r1_tarski(sK0,k1_zfmisc_1(k3_tarski(sK0))) <- (~1) [avatar component clause 68] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. ~1 [avatar split clause 46,68]
113. ~r1_tarski(sK1(X1,k1_zfmisc_1(X2)),X2) | r1_tarski(X1,k1_zfmisc_1(X2)) [resolution 54,65] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
129. ~r2_hidden(sK1(X1,k1_zfmisc_1(k3_tarski(X2))),X2) | r1_tarski(X1,k1_zfmisc_1(k3_tarski(X2))) [resolution 113,51] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
179. r1_tarski(X0,k1_zfmisc_1(k3_tarski(X0))) | r1_tarski(X0,k1_zfmisc_1(k3_tarski(X0))) [resolution 129,53] {a:3,w:10,wCS:4000,sine:1,thAx:0,allAx:7,thDist:-7}
183. r1_tarski(X0,k1_zfmisc_1(k3_tarski(X0))) [duplicate literal removal 179] {a:3,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:7,thDist:-7}
184. $false <- (~1) [unit resulting resolution 70,183] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
191. 1 [avatar contradiction clause 184]
192. $false [avatar sat refutation 71,191] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:40,thDist:-40}
% SZS output end Proof for zfmisc_1__t82_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 114
% Active clauses: 48
% Passive clauses: 73
% Final active clauses: 47
% Final passive clauses: 24
% Main loop iterations started: 49
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Equational tautologies: 2
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 73
% Unit resulting resolution: 1
% Factoring: 1
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 8
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% sat proof minimization: 0.001 s
% term sharing: 0.001 s
% non unit clause index maintenance: 0.001 s ( own 0.0000 s ) 
% other: 0.002 s ( own 0.0000 s ) 

