i: [1,284,0,2,1,1,0,0] t215_xreal_1
a: [3,451,0,2,1,284]
s: 451
i: [1,285,0,3,1,1,0,0] t215_xreal_1
a: [3,446,0,3,1,285]
s: 446
i: [1,286,0,5,1,1,0,0] t215_xreal_1
a: [3,441,0,5,1,286]
s: 441
i: [1,287,0,2,1,0,0,2] fc1_numbers
a: [3,456,0,2,1,287]
s: 456
i: [1,288,0,2,1,0,0,1] spc2_boole
a: [3,461,0,2,1,288]
s: 461
i: [1,289,0,2,1,0,0,2] spc1_boole
a: [3,466,0,2,1,289]
s: 466
i: [1,290,0,2,1,0,0,255] fc1_xboole_0
a: [3,471,0,2,1,290]
s: 471
i: [1,318,0,2,1,0,0,255] fc6_ordinal1
a: [3,601,0,2,1,318]
s: 601
i: [1,317,0,2,1,0,0,255] fc6_ordinal1
a: [3,606,0,2,1,317]
s: 606
i: [1,319,0,2,1,0,0,2] spc2_numerals
a: [3,626,0,2,1,319]
s: 626
i: [1,323,0,2,1,0,0,2] spc1_numerals
a: [3,646,0,2,1,323]
s: 646
i: [1,405,0,2,1,0,0,2] rc2_xboole_0
a: [3,658,0,2,1,405]
s: 658
i: [1,406,0,2,1,0,0,255] rc4_ordinal1
a: [3,663,0,2,1,406]
s: 663
i: [1,407,0,2,1,0,0,255] rc1_ordinal1
a: [3,668,0,2,1,407]
s: 668
i: [1,408,0,2,1,0,0,1] rc1_xreal_0
a: [3,673,0,2,1,408]
s: 673
i: [1,409,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,678,0,2,1,409]
s: 678
i: [1,410,0,2,1,0,0,255] rc1_xxreal_0
a: [3,683,0,2,1,410]
s: 683
i: [1,411,0,2,1,0,0,2] rc1_xboole_0
a: [3,688,0,2,1,411]
s: 688
i: [1,415,0,2,1,0,0,255] rc2_ordinal1
a: [3,693,0,2,1,415]
s: 693
i: [1,414,0,2,1,0,0,255] rc2_ordinal1
a: [3,698,0,2,1,414]
s: 698
i: [1,413,0,2,1,0,0,255] rc2_ordinal1
a: [3,703,0,2,1,413]
s: 703
i: [1,412,0,2,1,0,0,255] rc2_ordinal1
a: [3,708,0,2,1,412]
s: 708
i: [1,417,0,2,1,0,0,255] rc5_ordinal1
a: [3,713,0,2,1,417]
s: 713
i: [1,416,0,2,1,0,0,255] rc5_ordinal1
a: [3,718,0,2,1,416]
s: 718
i: [1,419,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,723,0,2,1,419]
s: 723
i: [1,418,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,728,0,2,1,418]
s: 728
i: [1,423,0,2,1,0,0,255] rc3_xreal_0
a: [3,733,0,2,1,423]
s: 733
i: [1,422,0,2,1,0,0,255] rc3_xreal_0
a: [3,738,0,2,1,422]
s: 738
i: [1,421,0,2,1,0,0,255] rc3_xreal_0
a: [3,743,0,2,1,421]
s: 743
i: [1,420,0,2,1,0,0,255] rc3_xreal_0
a: [3,748,0,2,1,420]
s: 748
i: [1,427,0,2,1,0,0,255] rc2_xreal_0
a: [3,753,0,2,1,427]
s: 753
i: [1,426,0,2,1,0,0,255] rc2_xreal_0
a: [3,758,0,2,1,426]
s: 758
i: [1,425,0,2,1,0,0,255] rc2_xreal_0
a: [3,763,0,2,1,425]
s: 763
i: [1,424,0,2,1,0,0,255] rc2_xreal_0
a: [3,768,0,2,1,424]
s: 768
i: [1,429,0,2,1,0,0,255] rc3_xxreal_0
a: [3,773,0,2,1,429]
s: 773
i: [1,428,0,2,1,0,0,255] rc3_xxreal_0
a: [3,778,0,2,1,428]
s: 778
i: [1,431,0,2,1,0,0,255] rc2_xxreal_0
a: [3,783,0,2,1,431]
s: 783
i: [1,430,0,2,1,0,0,255] rc2_xxreal_0
a: [3,788,0,2,1,430]
s: 788
i: [1,435,0,2,1,0,0,255] rc4_xreal_0
a: [3,793,0,2,1,435]
s: 793
i: [1,434,0,2,1,0,0,255] rc4_xreal_0
a: [3,798,0,2,1,434]
s: 798
i: [1,433,0,2,1,0,0,255] rc4_xreal_0
a: [3,803,0,2,1,433]
s: 803
i: [1,432,0,2,1,0,0,255] rc4_xreal_0
a: [3,808,0,2,1,432]
s: 808
i: [1,437,0,2,1,0,0,255] rc4_xxreal_0
a: [3,813,0,2,1,437]
s: 813
i: [1,436,0,2,1,0,0,255] rc4_xxreal_0
a: [3,818,0,2,1,436]
s: 818
i: [1,380,0,3,1,0,0,255] reflexivity_r1_tarski
s: 380
i: [1,291,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r2_r1
a: [3,476,0,3,1,291]
s: 476
i: [1,292,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,481,0,3,1,292]
s: 481
i: [1,293,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,486,0,3,1,293]
s: 486
r: 471
d: [2,820,0,2,1,2,23,471,486]
a: [3,824,0,2,1,820]
s: 824
i: [1,294,0,3,1,0,0,255] t1_numerals
a: [3,491,0,3,1,294]
s: 491
i: [1,295,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r2_r2
a: [3,496,0,3,1,295]
s: 496
i: [1,296,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r2
a: [3,501,0,3,1,296]
s: 501
i: [1,297,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,506,0,3,1,297]
s: 506
i: [1,322,0,3,1,0,0,2] spc2_numerals
a: [3,611,0,3,1,322]
s: 611
i: [1,321,0,3,1,0,0,2] spc2_numerals
a: [3,616,0,3,1,321]
s: 616
i: [1,326,0,3,1,0,0,2] spc1_numerals
a: [3,631,0,3,1,326]
s: 631
i: [1,325,0,3,1,0,0,2] spc1_numerals
a: [3,636,0,3,1,325]
s: 636
d: [2,826,0,3,1,2,22,491,486]
a: [3,830,0,3,1,826]
s: 830
d: [2,832,0,3,1,2,22,616,481]
a: [3,836,0,3,1,832]
s: 836
d: [2,838,0,3,1,2,22,636,481]
a: [3,842,0,3,1,838]
s: 842
i: [1,328,0,4,2,0,0,255] cc6_ordinal1
s: 328
d: [2,844,1,2,1,1,40,328,663]
a: [3,849,1,2,1,844]
s: 849
i: [1,329,0,4,2,0,0,255] cc2_xreal_0
s: 329
d: [2,845,1,2,1,1,40,328,713]
a: [3,854,1,2,1,845]
s: 854
i: [1,330,0,4,2,0,0,255] cc2_xcmplx_0
s: 330
i: [1,331,0,4,2,0,0,255] cc2_xxreal_0
s: 331
d: [2,856,1,2,1,1,40,329,663]
a: [3,861,1,2,1,856]
s: 861
i: [1,332,0,4,2,0,0,255] cc1_ordinal1
s: 332
d: [2,857,1,2,1,1,40,329,713]
a: [3,866,1,2,1,857]
s: 866
i: [1,333,0,4,2,0,0,255] cc1_ordinal1
s: 333
i: [1,335,0,4,2,0,0,2] cc3_xreal_0
s: 335
d: [2,942,1,2,1,1,40,335,451]
a: [3,952,1,2,1,942]
s: 952
i: [1,336,0,4,2,0,0,255] cc4_xreal_0
s: 336
d: [2,961,1,2,1,1,40,336,451]
a: [3,971,1,2,1,961]
s: 971
i: [1,349,0,4,2,0,0,255] cc9_ordinal1
s: 349
i: [1,350,0,4,2,0,0,255] cc7_ordinal1
s: 350
d: [2,868,1,2,1,1,40,330,663]
a: [3,873,1,2,1,868]
s: 873
i: [1,351,0,4,2,0,0,255] cc3_ordinal1
s: 351
d: [2,869,1,2,1,1,40,330,713]
a: [3,878,1,2,1,869]
s: 878
i: [1,379,0,4,1,0,0,3] existence_m1_subset_1
s: 379
i: [1,298,0,4,1,0,0,2] dt_k5_numbers
a: [3,511,0,4,1,298]
s: 511
d: [2,880,1,2,1,1,40,331,663]
a: [3,885,1,2,1,880]
s: 885
i: [1,299,0,4,1,0,0,1] dt_k6_numbers
a: [3,516,0,4,1,299]
s: 516
d: [2,881,1,2,1,1,40,331,713]
a: [3,890,1,2,1,881]
s: 890
i: [1,320,0,4,1,0,0,2] spc2_numerals
a: [3,621,0,4,1,320]
s: 621
i: [1,324,0,4,1,0,0,2] spc1_numerals
a: [3,641,0,4,1,324]
s: 641
d: [2,892,1,2,1,1,40,332,601]
a: [3,900,1,2,1,892]
s: 900
d: [2,1052,0,4,1,2,22,511,481]
a: [3,1056,0,4,1,1052]
s: 1056
d: [2,894,1,2,1,1,40,332,668]
a: [3,910,1,2,1,894]
s: 910
d: [2,1058,0,4,1,2,22,516,481]
a: [3,1062,0,4,1,1058]
s: 1062
d: [2,1064,0,4,1,2,22,621,481]
a: [3,1068,0,4,1,1064]
s: 1068
d: [2,917,1,2,1,1,40,333,601]
a: [3,925,1,2,1,917]
s: 925
d: [2,1070,0,4,1,2,22,641,481]
a: [3,1074,0,4,1,1070]
s: 1074
d: [2,919,1,2,1,1,40,333,668]
a: [3,935,1,2,1,919]
s: 935
i: [1,352,0,5,2,0,0,255] t6_boole
s: 352
i: [1,353,0,5,2,0,0,255] cc8_ordinal1
s: 353
d: [2,944,1,2,1,1,40,335,673]
a: [3,958,1,2,1,944]
s: 958
i: [1,354,0,5,2,0,0,2] cc1_xreal_0
s: 354
d: [2,963,1,2,1,1,40,336,673]
a: [3,977,1,2,1,963]
s: 977
i: [1,355,0,5,2,0,0,2] cc1_xcmplx_0
s: 355
i: [1,376,0,5,2,0,0,255] t7_boole
s: 376
d: [2,980,1,2,1,1,40,349,824]
a: [3,987,1,2,1,980]
s: 987
i: [1,300,0,5,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r2_rn1d2
a: [3,521,0,5,1,300]
s: 521
d: [2,981,1,2,1,1,40,349,688]
a: [3,992,1,2,1,981]
s: 992
i: [1,301,0,5,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_rn1d2
a: [3,526,0,5,1,301]
s: 526
i: [1,302,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r2_r1_r2
a: [3,531,0,5,1,302]
s: 531
d: [2,982,1,2,1,1,40,349,808]
a: [3,997,1,2,1,982]
s: 997
i: [1,303,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_r2_r2
a: [3,536,0,5,1,303]
s: 536
d: [2,983,1,2,1,1,40,349,818]
a: [3,1002,1,2,1,983]
s: 1002
i: [1,304,0,5,1,0,0,1] rqRealDiv__k7_xcmplx_0__r2_r1_r2
a: [3,541,0,5,1,304]
s: 541
i: [1,305,0,5,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,546,0,5,1,305]
s: 546
d: [2,1004,1,2,1,1,40,350,824]
a: [3,1011,1,2,1,1004]
s: 1011
i: [1,306,0,5,1,0,0,1] rqRealDiv__k7_xcmplx_0__r2_r2_r1
a: [3,551,0,5,1,306]
s: 551
d: [2,1005,1,2,1,1,40,350,688]
a: [3,1016,1,2,1,1005]
s: 1016
i: [1,307,0,5,1,0,0,2] rqRealDiv__k7_xcmplx_0__r1_r1_r1
a: [3,556,0,5,1,307]
s: 556
i: [1,308,0,5,1,0,0,1] rqLessOrEqual__r1_xxreal_0__rn1d2_r2
a: [3,561,0,5,1,308]
s: 561
d: [2,1006,1,2,1,1,40,350,808]
a: [3,1021,1,2,1,1006]
s: 1021
i: [1,309,0,5,1,0,0,1] rqLessOrEqual__r1_xxreal_0__rn1d2_r1
a: [3,566,0,5,1,309]
s: 566
d: [2,1007,1,2,1,1,40,350,818]
a: [3,1026,1,2,1,1007]
s: 1026
i: [1,401,0,7,3,0,0,255] reflexivity_r1_xxreal_0
a: [3,653,0,5,2,401]
s: 653
d: [2,1076,0,5,2,1,22,352,486]
s: 1076
d: [2,1195,1,5,2,2,44,446,1076]
s: 1195
i: [1,360,0,6,3,0,0,255] cc2_ordinal1
s: 360
d: [2,1201,2,4,2,3,40,1195,653]
d: [2,1202,2,2,1,4,21,1201,971]
a: [3,1206,2,2,1,1202]
s: 1206
i: [1,363,0,6,3,0,0,255] cc5_xxreal_0
s: 363
i: [1,365,0,6,3,0,0,255] cc5_xxreal_0
s: 365
d: [2,1194,1,7,2,2,44,441,1076]
s: 1194
i: [1,366,0,6,3,0,0,255] cc3_xxreal_0
s: 366
d: [2,1028,1,2,1,1,40,351,824]
a: [3,1035,1,2,1,1028]
s: 1035
i: [1,368,0,6,3,0,0,255] cc3_xxreal_0
s: 368
i: [1,370,0,6,3,0,0,255] cc7_xxreal_0
s: 370
d: [2,1029,1,2,1,1,40,351,688]
a: [3,1040,1,2,1,1029]
s: 1040
i: [1,371,0,6,3,0,0,255] cc7_xxreal_0
s: 371
d: [2,1030,1,2,1,1,40,351,808]
a: [3,1045,1,2,1,1030]
s: 1045
i: [1,372,0,6,2,0,0,255] t1_subset
s: 372
i: [1,381,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 381
d: [2,1031,1,2,1,1,40,351,818]
a: [3,1050,1,2,1,1031]
s: 1050
i: [1,334,0,7,3,0,0,255] cc5_ordinal1
s: 334
d: [2,1078,1,2,1,1,40,353,836]
a: [3,1085,1,2,1,1078]
s: 1085
i: [1,345,0,7,2,0,0,1] t2_arithm
s: 345
i: [1,346,0,7,2,0,0,1] t5_arithm
s: 346
d: [2,1079,1,2,1,1,40,353,842]
a: [3,1090,1,2,1,1079]
s: 1090
i: [1,347,0,7,2,0,0,2] t3_arithm
s: 347
d: [2,1097,1,2,1,1,40,354,611]
a: [3,1103,1,2,1,1097]
s: 1103
i: [1,348,0,7,2,0,0,2] t6_arithm
s: 348
i: [1,374,0,7,2,0,0,255] t3_subset
s: 374
d: [2,1098,1,2,1,1,40,354,631]
a: [3,1108,1,2,1,1098]
s: 1108
i: [1,375,0,7,3,0,0,2] t8_boole
s: 375
d: [2,1115,1,2,1,1,40,355,611]
a: [3,1121,1,2,1,1115]
s: 1121
i: [1,310,0,7,1,0,0,1] rqRealDiv__k7_xcmplx_0__r1_rn1d2_r2
a: [3,571,0,7,1,310]
s: 571
i: [1,311,0,7,1,0,0,2] rqRealMult__k3_xcmplx_0__r2_rn1d2_r1
a: [3,576,0,7,1,311]
s: 576
d: [2,1116,1,2,1,1,40,355,631]
a: [3,1126,1,2,1,1116]
s: 1126
i: [1,312,0,7,1,0,0,2] rqRealMult__k3_xcmplx_0__rn1d2_r2_r1
a: [3,581,0,7,1,312]
s: 581
d: [2,1211,1,4,2,1,40,363,738]
d: [2,1213,1,2,1,2,21,1211,743]
a: [3,1217,1,2,1,1213]
s: 1217
i: [1,314,0,7,1,0,0,1] rqLessOrEqual__r1_xxreal_0__rn1d2_rn1d2
a: [3,586,0,7,1,314]
s: 586
i: [1,327,0,8,3,0,0,255] cc10_ordinal1
s: 327
d: [2,1212,1,4,2,1,40,363,773]
d: [2,1219,1,2,1,2,21,1212,778]
a: [3,1223,1,2,1,1219]
s: 1223
i: [1,357,0,8,4,0,0,255] cc4_xxreal_0
s: 357
d: [2,1388,2,6,3,1,40,357,971]
d: [2,1399,2,4,2,2,21,1388,1206]
a: [3,1403,2,2,1,1399]
s: 1403
i: [1,359,0,8,4,0,0,255] cc6_xxreal_0
s: 359
i: [1,361,0,8,4,0,0,255] cc8_xxreal_0
s: 361
a: [3,1406,2,2,1,1399]
s: 1406
i: [1,373,0,8,3,0,0,255] t2_subset
s: 373
d: [2,1225,1,4,2,1,40,365,626]
a: [3,1231,1,2,1,1225]
s: 1231
i: [1,396,0,8,3,0,0,2] fc6_xreal_0
s: 396
i: [1,397,0,8,3,0,0,1] fc8_xreal_0
s: 397
a: [3,1236,1,2,1,1225]
s: 1236
i: [1,398,0,8,3,0,0,2] fc3_xcmplx_0
s: 398
d: [2,1226,1,4,2,1,40,365,646]
a: [3,1240,1,2,1,1226]
s: 1240
i: [1,399,0,8,3,0,0,2] fc5_xcmplx_0
s: 399
i: [1,378,0,9,3,0,0,255] t5_subset
s: 378
a: [3,1245,1,2,1,1226]
s: 1245
i: [1,315,0,9,1,0,0,2] rqRealMult__k3_xcmplx_0__r1_rn1d2_rn1d2
a: [3,591,0,9,1,315]
s: 591
d: [2,1227,1,4,2,1,40,365,758]
d: [2,1247,1,2,1,2,21,1227,763]
a: [3,1251,1,2,1,1247]
s: 1251
i: [1,316,0,9,1,0,0,2] rqRealMult__k3_xcmplx_0__rn1d2_r1_rn1d2
a: [3,596,0,9,1,316]
s: 596
i: [1,377,0,10,3,0,0,255] t4_subset
s: 377
d: [2,1228,1,4,2,1,40,365,783]
d: [2,1253,1,2,1,2,21,1228,788]
a: [3,1257,1,2,1,1253]
s: 1257
i: [1,402,0,10,4,0,0,255] connectedness_r1_xxreal_0
s: 402
d: [2,1561,1,7,3,1,40,402,446]
d: [2,1568,1,5,2,2,21,1561,971]
a: [3,1572,1,3,1,1568]
s: 1572
i: [1,339,0,11,5,0,0,1] t2_real
s: 339
i: [1,340,0,11,5,0,0,1] t1_real
s: 340
d: [2,1580,2,14,5,2,40,339,1194]
a: [3,1596,2,4,1,1580]
s: 1596
d: [2,1640,3,0,0,3,52,451,1103,1596,397]
e: 1640
d: [2,1641,3,4,2,1,40,1596,397]
d: [2,1643,3,2,1,2,21,1641,1103]
d: [2,1644,3,0,0,3,21,1643,451]
e: 1644
r: 1596
i: [1,341,0,11,5,0,0,1] t8_real
s: 341
a: [3,1595,2,4,1,1580]
s: 1595
i: [1,342,0,11,5,0,0,1] t7_real
s: 342
i: [1,343,0,11,5,0,0,1] t3_real
s: 343
a: [3,1600,2,4,1,1580]
s: 1600
i: [1,344,0,11,5,0,0,1] t4_real
s: 344
d: [2,1614,2,14,5,2,40,340,1194]
a: [3,1633,2,4,1,1614]
s: 1633
i: [1,400,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 400
i: [1,382,0,12,5,0,0,2] fc26_xreal_0
s: 382
d: [2,1655,3,4,1,1,40,1595,336]
a: [3,1660,3,4,1,1655]
s: 1660
i: [1,383,0,12,5,0,0,2] fc33_xreal_0
s: 383
d: [2,1731,3,0,0,5,52,1406,451,1236,1103,1600,383]
e: 1731
d: [2,1732,3,8,4,1,40,383,1600]
d: [2,1742,3,6,3,2,21,1732,1103]
d: [2,1743,3,4,2,3,21,1742,1236]
d: [2,1744,3,2,1,4,21,1743,451]
d: [2,1745,3,0,0,5,21,1744,1406]
e: 1745
r: 1403
r: 1406
r: 1600
r: 1633
d: [2,1581,1,12,4,1,40,339,441]
a: [3,1608,1,2,1,1581]
s: 1608
i: [1,384,0,12,5,0,0,2] fc32_xreal_0
s: 384
i: [1,385,0,12,5,0,0,2] fc31_xreal_0
s: 385
d: [2,1615,1,12,4,1,40,340,441]
a: [3,1637,1,2,1,1615]
s: 1637
i: [1,390,0,12,5,0,0,2] fc8_xcmplx_0
s: 390
s: 1403
i: [1,391,0,12,5,0,0,2] fc9_xcmplx_0
s: 391
i: [1,392,0,12,5,0,0,2] fc24_xreal_0
s: 392
s: 1406
i: [1,393,0,12,5,0,0,2] fc23_xreal_0
s: 393
a: [3,1599,2,4,1,1580]
s: 1599
i: [1,394,0,12,5,0,0,2] fc25_xreal_0
s: 394
i: [1,395,0,12,5,0,0,2] fc34_xreal_0
s: 395
s: 1633
i: [1,337,0,13,6,0,0,1] t5_real
s: 337
d: [2,1762,2,4,2,2,44,1608,1076]
s: 1762
r: 363
i: [1,338,0,13,6,0,0,1] t6_real
s: 338
i: [1,389,0,14,4,0,0,2] existence_m2_subset_1
s: 389
d: [2,1797,2,4,2,2,44,1637,1076]
s: 1797
r: 366
i: [1,386,0,15,5,0,0,2] dt_m2_subset_1
s: 386
d: [2,1656,3,4,1,1,40,1595,335]
a: [3,1665,3,4,1,1656]
s: 1665
i: [1,387,0,15,5,0,0,2] redefinition_m2_subset_1
s: 387
i: [1,388,0,15,5,0,0,2] redefinition_m2_subset_1
s: 388
d: [2,1667,3,8,2,1,40,1600,363]
a: [3,1670,3,4,1,1667]
s: 1670
d: [2,1975,4,0,0,5,52,1206,952,461,1121,1670,391]
e: 1975
d: [2,1976,4,8,4,1,40,1670,391]
d: [2,1981,4,6,3,2,21,1976,1121]
d: [2,1982,4,4,2,3,21,1981,461]
d: [2,1983,4,2,1,4,21,1982,952]
d: [2,1984,4,0,0,5,21,1983,1206]
e: 1984
r: 1403
r: 1406
r: 1599
r: 1762
r: 1608
r: 1670
s: 363
a: [3,1607,1,2,1,1581]
s: 1607
r: 653
r: 1206
r: 1607
r: 1633
r: 1797
r: 1637
a: [3,650,0,2,1,401]
s: 650
r: 402
r: 365
r: 363
r: 357
d: [2,2040,1,0,0,2,40,650,813]
e: 2040
d: [2,2039,1,0,0,2,40,650,798]
e: 2039
d: [2,2038,1,0,0,2,40,650,788]
e: 2038
d: [2,2037,1,0,0,2,40,650,778]
e: 2037
d: [2,2036,1,0,0,2,40,650,763]
e: 2036
d: [2,2035,1,0,0,2,40,650,743]
e: 2035
d: [2,2034,2,0,0,2,40,650,890]
e: 2034
d: [2,2033,1,0,0,2,40,650,683]
e: 2033
d: [2,2032,2,0,0,2,40,650,977]
e: 2032
d: [2,2031,2,0,0,2,40,650,885]
e: 2031
d: [2,2030,2,0,0,2,40,650,971]
e: 2030
d: [2,2029,2,0,0,2,40,650,1240]
e: 2029
d: [2,2028,4,0,0,2,40,650,1660]
e: 2028
d: [2,2027,2,0,0,2,40,650,1231]
e: 2027
d: [2,2024,2,0,0,2,52,1231,650]
e: 2024
d: [2,2025,2,0,0,2,52,1103,650,336]
e: 2025
d: [2,2026,2,0,0,2,52,1011,650,331]
e: 2026
r: 650
s: 653
s: 1608
s: 363
a: [3,1638,1,2,1,1615]
s: 1638
s: 365
s: 366
d: [2,1077,1,2,1,1,40,353,830]
d: [2,1133,2,2,1,1,40,1011,331]
d: [2,1135,2,2,1,1,40,1011,329]
d: [2,1367,2,2,1,1,40,1103,336]
d: [2,1651,1,8,4,1,40,341,446]
d: [2,1752,1,6,3,2,21,1651,451]
d: [2,1815,3,4,2,1,40,1403,365]
d: [2,1816,3,2,1,2,21,1815,971]
d: [2,2005,2,4,2,1,40,1607,363]
d: [2,2007,2,2,1,2,21,2005,824]
d: [2,1566,1,7,3,0,42,402]
d: [2,1567,1,5,2,0,18,1566]
d: [2,1925,1,16,5,1,40,338,441]
d: [2,2020,1,12,4,2,21,1925,1595]
d: [2,2062,2,4,2,3,40,653,1195]
d: [2,2071,2,2,1,4,21,2062,971]
d: [2,2082,2,4,2,1,40,366,1638]
d: [2,2087,2,2,1,2,21,2082,824]
f: 2087,2071,2020,1567,2007,1984,1816,1744,1752,1644,1367,1225,1135,1133,1115,1097,1077,961,942,826,820,319,322,294,293,290,288,284,285,286
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t215_xreal_1
% SZS output start Proof for xreal_1__t215_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ~(r1_xxreal_0(k7_xcmplx_0(X0,2),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers))) [input t215_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ~(r1_xxreal_0(k7_xcmplx_0(X0,2),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers))) [negated conjecture 1]
12. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
13. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
16. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
18. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
25. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc33_xreal_0]
32. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
40. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0))))) [input t6_real]
42. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
52. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xcmplx_0(X0)) [input cc1_xcmplx_0]
53. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xreal_0(X0)) [input cc1_xreal_0]
55. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
62. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
69. ! [X0,X1] : ((v1_xcmplx_0(X1) & ~v1_xboole_0(X1) & v1_xcmplx_0(X0) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k7_xcmplx_0(X0,X1))) [input fc9_xcmplx_0]
88. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
93. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
94. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
98. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
99. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
101. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
123. ~v1_xboole_0(2) [input spc2_boole]
125. m1_subset_1(2,k1_numbers) & m1_subset_1(2,k5_numbers) & m2_subset_1(2,k1_numbers,k5_numbers) & v2_xxreal_0(2) [input spc2_numerals]
139. ? [X0] : ((r1_xxreal_0(k7_xcmplx_0(X0,2),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers)) & v1_xreal_0(X0)) [ennf transformation 2]
140. ? [X0] : (r1_xxreal_0(k7_xcmplx_0(X0,2),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X0)) [flattening 139]
143. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
145. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 13]
148. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 98]
149. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 99]
152. ! [X0] : (! [X1] : (((v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0)) | ~r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 40]
153. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 152]
158. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 42]
159. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 158]
173. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 88]
174. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 18]
175. ! [X0] : (v1_xreal_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 53]
176. ! [X0] : (v1_xcmplx_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 52]
185. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 16]
186. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 185]
187. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 55]
188. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 187]
203. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 25]
204. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 203]
217. ! [X0,X1] : (~v1_xboole_0(k7_xcmplx_0(X0,X1)) | (~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0))) [ennf transformation 69]
218. ! [X0,X1] : (~v1_xboole_0(k7_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0)) [flattening 217]
229. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 101]
230. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 229]
239. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 93]
240. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 239]
245. ? [X0] : (r1_xxreal_0(k7_xcmplx_0(X0,2),k6_numbers) & ~r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X0)) => (r1_xxreal_0(k7_xcmplx_0(sK0,2),k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK0)) [choice axiom]
246. r1_xxreal_0(k7_xcmplx_0(sK0,2),k6_numbers) & ~r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK0) [skolemisation 140,245]
284. v1_xreal_0(sK0) [cnf transformation 246] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
285. ~r1_xxreal_0(sK0,k6_numbers) [cnf transformation 246] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
286. r1_xxreal_0(k7_xcmplx_0(sK0,2),k6_numbers) [cnf transformation 246] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
288. ~v1_xboole_0(2) [cnf transformation 123] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
290. v1_xboole_0(k1_xboole_0) [cnf transformation 62] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
293. k6_numbers = k1_xboole_0 [cnf transformation 94] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
294. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 32] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
319. v2_xxreal_0(2) [cnf transformation 125] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
322. m1_subset_1(2,k1_numbers) [cnf transformation 125] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
329. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 143] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
331. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 145] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
335. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 148] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
336. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 149] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
338. ~r1_xxreal_0(X0,X1) | v2_xxreal_0(X1) | v1_xboole_0(X0) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 153] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
341. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 159] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
352. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 173] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
353. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 174] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
354. ~m1_subset_1(X0,k1_numbers) | v1_xreal_0(X0) [cnf transformation 175] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
355. ~m1_subset_1(X0,k1_numbers) | v1_xcmplx_0(X0) [cnf transformation 176] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
363. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 186] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
365. ~v2_xxreal_0(X0) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 186] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
366. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 188] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
383. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 204] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
391. ~v1_xboole_0(k7_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) [cnf transformation 218] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
397. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 230] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
402. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 240] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
439. 1 <=> r1_xxreal_0(k7_xcmplx_0(sK0,2),k6_numbers) [avatar definition]
441. r1_xxreal_0(k7_xcmplx_0(sK0,2),k6_numbers) <- (1) [avatar component clause 439] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
442. 1 [avatar split clause 286,439]
444. 2 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
446. ~r1_xxreal_0(sK0,k6_numbers) <- (~2) [avatar component clause 444] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
447. ~2 [avatar split clause 285,444]
449. 3 <=> v1_xreal_0(sK0) [avatar definition]
451. v1_xreal_0(sK0) <- (3) [avatar component clause 449] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
452. 3 [avatar split clause 284,449]
459. 5 <=> v1_xboole_0(2) [avatar definition]
461. ~v1_xboole_0(2) <- (~5) [avatar component clause 459] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
462. ~5 [avatar split clause 288,459]
469. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
471. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 469] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
472. 7 [avatar split clause 290,469]
484. 10 <=> k6_numbers = k1_xboole_0 [avatar definition]
486. k6_numbers = k1_xboole_0 <- (10) [avatar component clause 484] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
487. 10 [avatar split clause 293,484]
489. 11 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
491. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (11) [avatar component clause 489] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
492. 11 [avatar split clause 294,489]
609. 35 <=> m1_subset_1(2,k1_numbers) [avatar definition]
611. m1_subset_1(2,k1_numbers) <- (35) [avatar component clause 609] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
612. 35 [avatar split clause 322,609]
624. 38 <=> v2_xxreal_0(2) [avatar definition]
626. v2_xxreal_0(2) <- (38) [avatar component clause 624] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
627. 38 [avatar split clause 319,624]
652. 44 <=> ! [X0] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0)) [avatar definition]
653. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) <- (44) [avatar component clause 652] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
820. v1_xboole_0(k6_numbers) <- (7, 10) [backward demodulation 471,486] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
822. 78 <=> v1_xboole_0(k6_numbers) [avatar definition]
824. v1_xboole_0(k6_numbers) <- (78) [avatar component clause 822] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
825. 78 | ~7 | ~10 [avatar split clause 820,484,469,822]
826. m1_subset_1(k6_numbers,k4_ordinal1) <- (10, 11) [forward demodulation 491,486] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
828. 79 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
830. m1_subset_1(k6_numbers,k4_ordinal1) <- (79) [avatar component clause 828] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
831. 79 | ~10 | ~11 [avatar split clause 826,489,484,828]
942. v1_xcmplx_0(sK0) <- (3) [resolution 335,451] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
950. 98 <=> v1_xcmplx_0(sK0) [avatar definition]
952. v1_xcmplx_0(sK0) <- (98) [avatar component clause 950] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
953. 98 | ~3 [avatar split clause 942,449,950]
961. v1_xxreal_0(sK0) <- (3) [resolution 336,451] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
969. 100 <=> v1_xxreal_0(sK0) [avatar definition]
971. v1_xxreal_0(sK0) <- (100) [avatar component clause 969] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
972. 100 | ~3 [avatar split clause 961,449,969]
1009. 106 <=> v7_ordinal1(k6_numbers) [avatar definition]
1011. v7_ordinal1(k6_numbers) <- (106) [avatar component clause 1009] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1076. k6_numbers = X0 | ~v1_xboole_0(X0) <- (10) [forward demodulation 352,486] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1077. v7_ordinal1(k6_numbers) <- (79) [resolution 353,830] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1081. 106 | ~79 [avatar split clause 1077,828,1009]
1097. v1_xreal_0(2) <- (35) [resolution 354,611] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
1101. 121 <=> v1_xreal_0(2) [avatar definition]
1103. v1_xreal_0(2) <- (121) [avatar component clause 1101] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1104. 121 | ~35 [avatar split clause 1097,609,1101]
1115. v1_xcmplx_0(2) <- (35) [resolution 355,611] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
1119. 124 <=> v1_xcmplx_0(2) [avatar definition]
1121. v1_xcmplx_0(2) <- (124) [avatar component clause 1119] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1122. 124 | ~35 [avatar split clause 1115,609,1119]
1133. v1_xxreal_0(k6_numbers) <- (106) [resolution 1011,331] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1135. v1_xreal_0(k6_numbers) <- (106) [resolution 1011,329] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1138. 127 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1141. 127 | ~106 [avatar split clause 1133,1009,1138]
1148. 129 <=> v1_xreal_0(k6_numbers) [avatar definition]
1151. 129 | ~106 [avatar split clause 1135,1009,1148]
1195. ~r1_xxreal_0(sK0,X0) | ~v1_xboole_0(X0) <- (~2, 10) [superposition 446,1076] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1204. 135 <=> v1_xboole_0(sK0) [avatar definition]
1206. ~v1_xboole_0(sK0) <- (~135) [avatar component clause 1204] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1225. ~v3_xxreal_0(2) | ~v1_xxreal_0(2) <- (38) [resolution 365,626] {a:1,w:4,wCS:1600,sine:255,thAx:0,allAx:2,thDist:-2}
1230. 138 <=> v1_xxreal_0(2) [avatar definition]
1234. 139 <=> v3_xxreal_0(2) [avatar definition]
1236. ~v3_xxreal_0(2) <- (~139) [avatar component clause 1234] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1237. ~138 | ~139 | ~38 [avatar split clause 1225,624,1234,1230]
1367. v1_xxreal_0(2) <- (121) [resolution 1103,336] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1369. 138 | ~121 [avatar split clause 1367,1101,1230]
1401. 159 <=> v2_xxreal_0(sK0) [avatar definition]
1403. v2_xxreal_0(sK0) <- (159) [avatar component clause 1401] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1405. 160 <=> v3_xxreal_0(sK0) [avatar definition]
1566. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X0) [factoring 402] {a:1,w:7,wCS:2800,sine:255,thAx:0,allAx:1,thDist:-1}
1567. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) [duplicate literal removal 1566] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
1594. 180 <=> v1_xreal_0(k7_xcmplx_0(sK0,2)) [avatar definition]
1595. v1_xreal_0(k7_xcmplx_0(sK0,2)) <- (180) [avatar component clause 1594] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1596. ~v1_xreal_0(k7_xcmplx_0(sK0,2)) <- (~180) [avatar component clause 1594] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1598. 181 <=> v3_xxreal_0(k7_xcmplx_0(sK0,2)) [avatar definition]
1600. v3_xxreal_0(k7_xcmplx_0(sK0,2)) <- (181) [avatar component clause 1598] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1606. 183 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1607. v3_xxreal_0(k6_numbers) <- (183) [avatar component clause 1606] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1636. 186 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1638. v2_xxreal_0(k6_numbers) <- (186) [avatar component clause 1636] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1641. ~v1_xreal_0(2) | ~v1_xreal_0(sK0) <- (~180) [resolution 1596,397] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1643. ~v1_xreal_0(sK0) <- (121, ~180) [subsumption resolution 1641,1103] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1644. $false <- (3, 121, ~180) [subsumption resolution 1643,451] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1645. ~3 | ~121 | 180 [avatar contradiction clause 1644]
1651. v3_xxreal_0(k6_numbers) | v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (~2) [resolution 341,446] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1669. 189 <=> v1_xboole_0(k7_xcmplx_0(sK0,2)) [avatar definition]
1670. v1_xboole_0(k7_xcmplx_0(sK0,2)) <- (189) [avatar component clause 1669] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1732. ~v1_xreal_0(2) | v3_xxreal_0(2) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (181) [resolution 383,1600] {a:3,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1742. v3_xxreal_0(2) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (121, 181) [subsumption resolution 1732,1103] {a:3,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1743. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (121, ~139, 181) [subsumption resolution 1742,1236] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1744. v3_xxreal_0(sK0) <- (3, 121, ~139, 181) [subsumption resolution 1743,451] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1752. v3_xxreal_0(k6_numbers) | v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (~2, 3) [subsumption resolution 1651,451] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1753. ~129 | 159 | 183 | 2 | ~3 [avatar split clause 1752,449,444,1606,1401,1148]
1756. 160 | ~3 | ~121 | 139 | ~181 [avatar split clause 1744,1598,1234,1101,449,1405]
1815. ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (159) [resolution 1403,365] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1816. ~v3_xxreal_0(sK0) <- (100, 159) [subsumption resolution 1815,971] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1817. ~160 | ~100 | ~159 [avatar split clause 1816,1401,969,1405]
1925. v2_xxreal_0(k6_numbers) | v1_xboole_0(k7_xcmplx_0(sK0,2)) | v3_xxreal_0(k7_xcmplx_0(sK0,2)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,2)) <- (1) [resolution 338,441] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1976. ~v1_xcmplx_0(2) | v1_xboole_0(2) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (189) [resolution 1670,391] {a:4,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1981. v1_xboole_0(2) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (124, 189) [subsumption resolution 1976,1121] {a:4,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1982. ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (~5, 124, 189) [subsumption resolution 1981,461] {a:4,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1983. v1_xboole_0(sK0) <- (~5, 98, 124, 189) [subsumption resolution 1982,952] {a:4,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1984. $false <- (~5, 98, 124, ~135, 189) [subsumption resolution 1983,1206] {a:4,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1985. 5 | ~98 | ~124 | 135 | ~189 [avatar contradiction clause 1984]
2005. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (183) [resolution 1607,363] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2007. ~v1_xxreal_0(k6_numbers) <- (78, 183) [subsumption resolution 2005,824] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2008. ~127 | ~78 | ~183 [avatar split clause 2007,1606,822,1138]
2009. 44 [avatar split clause 1567,652]
2020. v2_xxreal_0(k6_numbers) | v1_xboole_0(k7_xcmplx_0(sK0,2)) | v3_xxreal_0(k7_xcmplx_0(sK0,2)) | ~v1_xreal_0(k6_numbers) <- (1, 180) [subsumption resolution 1925,1595] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
2021. ~129 | 181 | 189 | 186 | ~1 | ~180 [avatar split clause 2020,1594,439,1636,1669,1598,1148]
2062. ~v1_xxreal_0(sK0) | ~v1_xboole_0(sK0) <- (~2, 10, 44) [resolution 653,1195] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2071. ~v1_xboole_0(sK0) <- (~2, 10, 44, 100) [subsumption resolution 2062,971] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2072. ~135 | 2 | ~10 | ~44 | ~100 [avatar split clause 2071,969,652,484,444,1204]
2082. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (186) [resolution 366,1638] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2087. ~v1_xxreal_0(k6_numbers) <- (78, 186) [subsumption resolution 2082,824] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
2088. ~127 | ~78 | ~186 [avatar split clause 2087,1636,822,1138]
2093. $false [avatar sat refutation 442,447,452,462,472,487,492,612,627,825,831,953,972,1081,1104,1122,1141,1151,1237,1369,1645,1753,1756,1817,1985,2008,2009,2021,2072,2088] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1289,thDist:-1289}
% SZS output end Proof for xreal_1__t215_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 126
% >>> Preprocessing
% Introduced skolems: 19
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 154
% Generated clauses: 994
% Active clauses: 210
% Passive clauses: 331
% Final active clauses: 185
% Final passive clauses: 93
% Main loop iterations started: 226
% 
% >>> Simplifying Inferences
% Duplicate literals: 89
% Fw subsumption resolutions: 90
% Fw demodulations: 8
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 77
% Equational tautologies: 1
% Forward subsumptions: 167
% Backward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 270
% Unit resulting resolution: 6
% Factoring: 1
% Forward superposition: 130
% Backward superposition: 25
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 51
% Split components: 139
% Sat splitting refutations: 23
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 267
% SAT solver unit clauses: 79
% SAT solver binary clauses: 103
% 
% Memory used [KB]: 6268
% Time elapsed: 0.011 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% SAT solver time: 0.001 s
% congruence closure: 0.001 s
% forward superposition index maintenance: 0.001 s
% forward subsumption: 0.004 s
% other: 0.011 s ( own 0.003 s ) 

