i: [1,68,0,2,1,1,0,0] t30_pboole
a: [3,113,0,2,1,68]
s: 113
i: [1,66,0,2,1,1,0,0] t30_pboole
a: [3,123,0,2,1,66]
s: 123
i: [1,64,0,2,1,1,0,0] t30_pboole
a: [3,133,0,2,1,64]
s: 133
i: [1,62,0,2,1,1,0,0] t30_pboole
a: [3,143,0,2,1,62]
s: 143
i: [1,69,0,3,1,1,0,0] t30_pboole
a: [3,108,0,3,1,69]
s: 108
i: [1,67,0,3,1,1,0,0] t30_pboole
a: [3,118,0,3,1,67]
s: 118
i: [1,65,0,3,1,1,0,0] t30_pboole
a: [3,128,0,3,1,65]
s: 128
i: [1,63,0,3,1,1,0,0] t30_pboole
a: [3,138,0,3,1,63]
s: 138
i: [1,97,0,2,1,0,0,1] rc1_funct_1
a: [3,148,0,2,1,97]
s: 148
i: [1,96,0,2,1,0,0,1] rc1_funct_1
a: [3,153,0,2,1,96]
s: 153
i: [1,70,0,10,1,1,0,0] t30_pboole
a: [3,103,0,10,1,70]
s: 103
i: [1,75,0,3,1,0,0,1] rc1_pboole
s: 75
i: [1,77,0,3,1,0,0,1] rc1_pboole
s: 77
i: [1,76,0,4,1,0,0,1] rc1_pboole
s: 76
i: [1,78,0,4,1,0,0,1] rc1_pboole
s: 78
i: [1,98,0,24,9,0,0,1] redefinition_r6_pboole
d: [2,99,0,14,5,0,18,98]
s: 99
i: [1,79,0,24,9,0,0,1] reflexivity_r6_pboole
s: 79
i: [1,80,0,24,9,0,0,2] reflexivity_r2_pboole
s: 80
i: [1,85,0,25,9,0,0,1] dt_k2_pboole
s: 85
i: [1,87,0,25,9,0,0,1] dt_k2_pboole
s: 87
i: [1,89,0,25,9,0,0,1] dt_k3_pboole
s: 89
i: [1,91,0,25,9,0,0,1] dt_k3_pboole
s: 91
i: [1,81,0,26,9,0,0,1] idempotence_k2_pboole
s: 81
i: [1,82,0,26,9,0,0,1] idempotence_k3_pboole
s: 82
i: [1,86,0,26,9,0,0,1] dt_k2_pboole
s: 86
i: [1,88,0,26,9,0,0,1] dt_k2_pboole
s: 88
i: [1,90,0,26,9,0,0,1] dt_k3_pboole
s: 90
i: [1,92,0,26,9,0,0,1] dt_k3_pboole
s: 92
i: [1,71,0,27,9,0,0,1] t14_pboole
s: 71
i: [1,72,0,27,9,0,0,2] t15_pboole
s: 72
i: [1,94,0,27,10,0,0,1] redefinition_r6_pboole
s: 94
i: [1,93,0,28,10,0,0,1] symmetry_r6_pboole
s: 93
i: [1,83,0,29,9,0,0,1] commutativity_k2_pboole
s: 83
d: [2,184,1,47,17,0,44,71,83]
d: [2,193,1,27,9,0,18,184]
s: 193
i: [1,84,0,29,9,0,0,1] commutativity_k3_pboole
s: 84
i: [1,73,0,31,11,0,0,2] l22_pboole
s: 73
d: [2,179,1,47,17,0,44,71,83]
d: [2,198,1,27,9,0,18,179]
s: 198
i: [1,74,0,45,15,0,0,2] t17_pboole
s: 74
d: [2,210,1,47,17,0,44,72,84]
d: [2,219,1,27,9,0,18,210]
s: 219
d: [2,205,1,47,17,0,44,72,84]
d: [2,224,1,27,9,0,18,205]
s: 224
d: [2,226,1,65,18,0,40,73,71]
d: [2,231,1,55,14,0,18,226]
d: [2,233,1,49,13,0,21,231,88]
d: [2,234,1,44,12,0,21,233,87]
d: [2,235,1,38,11,0,21,234,86]
d: [2,236,1,33,10,0,21,235,85]
s: 236
d: [2,228,1,65,18,0,40,73,72]
d: [2,229,1,55,14,0,18,228]
d: [2,241,1,49,13,0,21,229,92]
d: [2,242,1,44,12,0,21,241,91]
d: [2,243,1,38,11,0,21,242,90]
d: [2,244,1,33,10,0,21,243,89]
s: 244
d: [2,245,1,83,24,0,40,74,73]
d: [2,252,1,73,20,0,18,245]
d: [2,253,1,67,19,0,21,252,92]
d: [2,254,1,62,18,0,21,253,91]
d: [2,255,1,56,17,0,21,254,90]
d: [2,256,1,51,16,0,21,255,89]
s: 256
d: [2,279,2,64,23,0,40,244,74]
d: [2,285,2,34,11,0,18,279]
d: [2,286,2,30,10,0,21,285,80]
s: 286
d: [2,287,2,64,23,0,40,256,72]
d: [2,301,2,34,11,0,18,287]
d: [2,302,2,30,10,0,21,301,80]
s: 302
d: [2,288,2,64,23,0,40,256,219]
d: [2,300,2,34,11,0,18,288]
d: [2,303,2,30,10,0,21,300,80]
s: 303
d: [2,330,3,43,10,1,44,103,286]
d: [2,354,3,40,9,2,21,330,128]
d: [2,355,3,38,8,3,21,354,133]
d: [2,356,3,35,7,4,21,355,138]
d: [2,357,3,33,6,5,21,356,143]
a: [3,365,3,5,1,357]
s: 365
d: [2,427,4,0,0,9,52,143,123,133,113,138,128,118,108,365,85]
e: 427
d: [2,428,4,20,8,1,40,365,85]
d: [2,430,4,17,7,2,21,428,108]
d: [2,431,4,15,6,3,21,430,113]
d: [2,432,4,12,5,4,21,431,118]
d: [2,433,4,10,4,5,21,432,123]
d: [2,434,4,7,3,6,21,433,128]
d: [2,435,4,5,2,7,21,434,133]
d: [2,436,4,2,1,8,21,435,138]
d: [2,437,4,0,0,9,21,436,143]
e: 437
r: 365
d: [2,227,2,65,18,0,40,73,193]
d: [2,230,2,55,14,0,18,227]
d: [2,237,2,49,13,0,21,230,88]
d: [2,238,2,44,12,0,21,237,87]
d: [2,239,2,38,11,0,21,238,86]
d: [2,240,2,33,10,0,21,239,85]
s: 240
a: [3,364,3,5,1,357]
s: 364
d: [2,257,2,65,18,0,40,219,73]
d: [2,264,2,55,14,0,18,257]
d: [2,265,2,49,13,0,21,264,92]
d: [2,266,2,44,12,0,21,265,91]
d: [2,267,2,38,11,0,21,266,90]
d: [2,268,2,33,10,0,21,267,89]
s: 268
d: [2,272,2,53,18,0,44,236,83]
d: [2,273,2,33,10,0,18,272]
s: 273
a: [3,369,3,6,1,357]
s: 369
d: [2,459,4,0,0,9,52,143,123,133,113,138,128,118,108,369,86]
e: 459
d: [2,460,4,20,8,1,40,369,86]
d: [2,462,4,17,7,2,21,460,108]
d: [2,463,4,15,6,3,21,462,113]
d: [2,464,4,12,5,4,21,463,118]
d: [2,465,4,10,4,5,21,464,123]
d: [2,466,4,7,3,6,21,465,128]
d: [2,467,4,5,2,7,21,466,133]
d: [2,468,4,2,1,8,21,467,138]
d: [2,469,4,0,0,9,21,468,143]
e: 469
r: 369
d: [2,271,2,53,18,0,44,236,83]
d: [2,274,2,33,10,0,18,271]
s: 274
a: [3,373,3,5,1,357]
s: 373
d: [2,471,4,0,0,9,52,143,123,133,113,138,128,118,108,373,87]
e: 471
d: [2,472,4,20,8,1,40,373,87]
d: [2,474,4,17,7,2,21,472,108]
d: [2,475,4,15,6,3,21,474,113]
d: [2,476,4,12,5,4,21,475,118]
d: [2,477,4,10,4,5,21,476,123]
d: [2,478,4,7,3,6,21,477,128]
d: [2,479,4,5,2,7,21,478,133]
d: [2,480,4,2,1,8,21,479,138]
d: [2,481,4,0,0,9,21,480,143]
e: 481
r: 373
d: [2,282,2,53,18,0,44,244,84]
d: [2,283,2,33,10,0,18,282]
s: 283
d: [2,281,2,53,18,0,44,244,84]
d: [2,284,2,33,10,0,18,281]
s: 284
a: [3,372,3,5,1,357]
s: 372
d: [2,290,2,118,23,0,40,256,71]
d: [2,304,2,109,22,0,21,290,88]
d: [2,305,2,101,21,0,21,304,87]
d: [2,306,2,92,20,0,21,305,86]
d: [2,307,2,84,19,0,21,306,85]
d: [2,308,2,78,18,0,21,307,92]
d: [2,309,2,73,17,0,21,308,91]
d: [2,310,2,67,16,0,21,309,90]
d: [2,311,2,62,15,0,21,310,89]
s: 311
a: [3,368,3,6,1,357]
s: 368
d: [2,291,2,118,23,0,40,256,193]
d: [2,312,2,109,22,0,21,291,88]
d: [2,313,2,101,21,0,21,312,87]
d: [2,314,2,92,20,0,21,313,86]
d: [2,315,2,84,19,0,21,314,85]
d: [2,316,2,78,18,0,21,315,92]
d: [2,317,2,73,17,0,21,316,91]
d: [2,318,2,67,16,0,21,317,90]
d: [2,319,2,62,15,0,21,318,89]
s: 319
d: [2,292,2,121,29,0,40,256,74]
d: [2,320,2,115,28,0,21,292,92]
d: [2,321,2,110,27,0,21,320,91]
d: [2,322,2,104,26,0,21,321,90]
d: [2,323,2,99,25,0,21,322,89]
d: [2,324,2,93,24,0,21,323,92]
d: [2,325,2,88,23,0,21,324,91]
d: [2,326,2,82,22,0,21,325,90]
d: [2,327,2,77,21,0,21,326,89]
s: 327
a: [3,377,3,6,1,357]
s: 377
d: [2,724,4,0,0,9,52,143,123,133,113,138,128,118,108,377,88]
e: 724
d: [2,725,4,20,8,1,40,377,88]
d: [2,727,4,17,7,2,21,725,108]
d: [2,728,4,15,6,3,21,727,113]
d: [2,729,4,12,5,4,21,728,118]
d: [2,730,4,10,4,5,21,729,123]
d: [2,731,4,7,3,6,21,730,128]
d: [2,732,4,5,2,7,21,731,133]
d: [2,733,4,2,1,8,21,732,138]
d: [2,734,4,0,0,9,21,733,143]
e: 734
r: 377
d: [2,269,2,118,23,0,40,236,74]
d: [2,275,2,112,22,0,21,269,92]
d: [2,276,2,107,21,0,21,275,91]
d: [2,277,2,101,20,0,21,276,90]
d: [2,278,2,96,19,0,21,277,89]
s: 278
a: [3,376,3,6,1,357]
s: 376
a: [3,361,3,7,1,357]
s: 361
d: [2,736,4,0,0,9,52,143,123,133,113,138,128,118,108,361,71]
e: 736
d: [2,737,4,20,8,1,40,361,71]
d: [2,739,4,17,7,2,21,737,108]
d: [2,740,4,15,6,3,21,739,113]
d: [2,741,4,12,5,4,21,740,118]
d: [2,742,4,10,4,5,21,741,123]
d: [2,743,4,7,3,6,21,742,128]
d: [2,744,4,5,2,7,21,743,133]
d: [2,745,4,2,1,8,21,744,138]
d: [2,746,4,0,0,9,21,745,143]
e: 746
r: 361
a: [3,381,3,4,1,357]
s: 381
d: [2,748,4,0,0,5,52,143,133,138,128,381,99]
e: 748
d: [2,749,4,10,4,1,40,381,99]
d: [2,751,4,7,3,2,21,749,128]
d: [2,752,4,5,2,3,21,751,133]
d: [2,753,4,2,1,4,21,752,138]
d: [2,754,4,0,0,5,21,753,143]
e: 754
f: 754,746,734,481,469,437,357,62,63,64,65,66,67,68,69,70
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pboole__t30_pboole
% SZS output start Proof for pboole__t30_pboole
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r6_pboole(X0,k3_pboole(X0,X1,k2_pboole(X0,X1,X2)),X1))) [input t30_pboole]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r6_pboole(X0,k3_pboole(X0,X1,k2_pboole(X0,X1,X2)),X1))) [negated conjecture 1]
5. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2)))) [input dt_k2_pboole]
6. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2)))) [input dt_k3_pboole]
9. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => X1 = X2))) [input l22_pboole]
12. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (r6_pboole(X0,X1,X2) <=> X1 = X2)) [input redefinition_r6_pboole]
13. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => r2_pboole(X0,X1,X1)) [input reflexivity_r2_pboole]
16. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,X1,k2_pboole(X0,X1,X2)))) [input t14_pboole]
17. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k3_pboole(X0,X1,X2),X1))) [input t15_pboole]
18. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X1,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,k3_pboole(X0,X2,X3)))))) [input t17_pboole]
19. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r6_pboole(X0,k3_pboole(X0,X1,k2_pboole(X0,X1,X2)),X1))) [flattening 2]
20. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,X1,k2_pboole(X0,X1,X2)))) [flattening 16]
21. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k3_pboole(X0,X1,X2),X1))) [flattening 17]
22. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => X1 = X2))) [flattening 9]
23. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X1,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,k3_pboole(X0,X2,X3)))))) [flattening 18]
24. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k3_pboole(X0,X1,k2_pboole(X0,X1,X2)),X1) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 19]
25. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k3_pboole(X0,X1,k2_pboole(X0,X1,X2)),X1) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 24]
26. ! [X0,X1] : (! [X2] : (r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 20]
27. ! [X0,X1] : (! [X2] : (r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 26]
28. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 21]
29. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 28]
30. ! [X0,X1] : (! [X2] : ((X1 = X2 | (~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 22]
31. ! [X0,X1] : (! [X2] : (X1 = X2 | ~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 30]
32. ! [X0,X1] : (! [X2] : (! [X3] : ((r2_pboole(X0,X1,k3_pboole(X0,X2,X3)) | (~r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 23]
33. ! [X0,X1] : (! [X2] : (! [X3] : (r2_pboole(X0,X1,k3_pboole(X0,X2,X3)) | ~r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 32]
36. ! [X0,X1,X2] : (r2_pboole(X0,X1,X1) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 13]
37. ! [X0,X1,X2] : (r2_pboole(X0,X1,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 36]
46. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 5]
47. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 46]
48. ! [X0,X1,X2] : ((v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 6]
49. ! [X0,X1,X2] : ((v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 48]
52. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 12]
53. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 52]
54. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k3_pboole(X0,X1,k2_pboole(X0,X1,X2)),X1) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (~r6_pboole(sK0,k3_pboole(sK0,sK1,k2_pboole(sK0,sK1,X2)),sK1) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
55. ? [X2] : (~r6_pboole(sK0,k3_pboole(sK0,sK1,k2_pboole(sK0,sK1,X2)),sK1) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (~r6_pboole(sK0,k3_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)),sK1) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
56. (~r6_pboole(sK0,k3_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)),sK1) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 25,55,54]
59. ! [X0,X1,X2] : (((r6_pboole(X0,X1,X2) | X1 != X2) & (X1 = X2 | ~r6_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 53]
62. v1_relat_1(sK1) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
63. v4_relat_1(sK1,sK0) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
64. v1_funct_1(sK1) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. v1_partfun1(sK1,sK0) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
66. v1_relat_1(sK2) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
67. v4_relat_1(sK2,sK0) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
68. v1_funct_1(sK2) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
69. v1_partfun1(sK2,sK0) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
70. ~r6_pboole(sK0,k3_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)),sK1) [cnf transformation 56] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
71. r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 27] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 29] {a:0,w:27,wCS:10800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
73. ~r2_pboole(X0,X2,X1) | X1 = X2 | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 31] {a:0,w:31,wCS:12400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
74. r2_pboole(X0,X1,k3_pboole(X0,X2,X3)) | ~r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 33] {a:0,w:45,wCS:18000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
80. r2_pboole(X0,X1,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 37] {a:0,w:24,wCS:9600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
85. v1_relat_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 47] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
86. v4_relat_1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 47] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. v1_funct_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 47] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
88. v1_partfun1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 47] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
89. v1_relat_1(k3_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 49] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
90. v4_relat_1(k3_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 49] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
91. v1_funct_1(k3_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 49] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
92. v1_partfun1(k3_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 49] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
95. r6_pboole(X0,X1,X2) | X1 != X2 | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 59] {a:0,w:27,wCS:10800,sine:1,thAx:0,allAx:1,thDist:-1}
98. r6_pboole(X0,X2,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) [equality resolution 95] {a:0,w:24,wCS:9600,sine:1,thAx:0,allAx:1,thDist:-1}
99. r6_pboole(X0,X2,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) [duplicate literal removal 98] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. 1 <=> r6_pboole(sK0,k3_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)),sK1) [avatar definition]
103. ~r6_pboole(sK0,k3_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)),sK1) <- (~1) [avatar component clause 101] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. ~1 [avatar split clause 70,101]
106. 2 <=> v1_partfun1(sK2,sK0) [avatar definition]
108. v1_partfun1(sK2,sK0) <- (2) [avatar component clause 106] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
109. 2 [avatar split clause 69,106]
111. 3 <=> v1_funct_1(sK2) [avatar definition]
113. v1_funct_1(sK2) <- (3) [avatar component clause 111] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 3 [avatar split clause 68,111]
116. 4 <=> v4_relat_1(sK2,sK0) [avatar definition]
118. v4_relat_1(sK2,sK0) <- (4) [avatar component clause 116] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. 4 [avatar split clause 67,116]
121. 5 <=> v1_relat_1(sK2) [avatar definition]
123. v1_relat_1(sK2) <- (5) [avatar component clause 121] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. 5 [avatar split clause 66,121]
126. 6 <=> v1_partfun1(sK1,sK0) [avatar definition]
128. v1_partfun1(sK1,sK0) <- (6) [avatar component clause 126] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. 6 [avatar split clause 65,126]
131. 7 <=> v1_funct_1(sK1) [avatar definition]
133. v1_funct_1(sK1) <- (7) [avatar component clause 131] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. 7 [avatar split clause 64,131]
136. 8 <=> v4_relat_1(sK1,sK0) [avatar definition]
138. v4_relat_1(sK1,sK0) <- (8) [avatar component clause 136] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
139. 8 [avatar split clause 63,136]
141. 9 <=> v1_relat_1(sK1) [avatar definition]
143. v1_relat_1(sK1) <- (9) [avatar component clause 141] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. 9 [avatar split clause 62,141]
228. k3_pboole(X10,X9,X11) = X9 | ~r2_pboole(X10,X9,k3_pboole(X10,X9,X11)) | ~v1_partfun1(k3_pboole(X10,X9,X11),X10) | ~v1_funct_1(k3_pboole(X10,X9,X11)) | ~v4_relat_1(k3_pboole(X10,X9,X11),X10) | ~v1_relat_1(k3_pboole(X10,X9,X11)) | ~v1_partfun1(X9,X10) | ~v1_funct_1(X9) | ~v4_relat_1(X9,X10) | ~v1_relat_1(X9) | ~v1_partfun1(X11,X10) | ~v1_funct_1(X11) | ~v4_relat_1(X11,X10) | ~v1_relat_1(X11) | ~v1_partfun1(X9,X10) | ~v1_funct_1(X9) | ~v4_relat_1(X9,X10) | ~v1_relat_1(X9) [resolution 73,72] {a:1,w:65,wCS:26000,sine:2,thAx:0,allAx:2,thDist:-2}
229. k3_pboole(X10,X9,X11) = X9 | ~r2_pboole(X10,X9,k3_pboole(X10,X9,X11)) | ~v1_partfun1(k3_pboole(X10,X9,X11),X10) | ~v1_funct_1(k3_pboole(X10,X9,X11)) | ~v4_relat_1(k3_pboole(X10,X9,X11),X10) | ~v1_relat_1(k3_pboole(X10,X9,X11)) | ~v1_partfun1(X9,X10) | ~v1_funct_1(X9) | ~v4_relat_1(X9,X10) | ~v1_relat_1(X9) | ~v1_partfun1(X11,X10) | ~v1_funct_1(X11) | ~v4_relat_1(X11,X10) | ~v1_relat_1(X11) [duplicate literal removal 228] {a:1,w:55,wCS:22000,sine:2,thAx:0,allAx:2,thDist:-2}
241. k3_pboole(X10,X9,X11) = X9 | ~r2_pboole(X10,X9,k3_pboole(X10,X9,X11)) | ~v1_funct_1(k3_pboole(X10,X9,X11)) | ~v4_relat_1(k3_pboole(X10,X9,X11),X10) | ~v1_relat_1(k3_pboole(X10,X9,X11)) | ~v1_partfun1(X9,X10) | ~v1_funct_1(X9) | ~v4_relat_1(X9,X10) | ~v1_relat_1(X9) | ~v1_partfun1(X11,X10) | ~v1_funct_1(X11) | ~v4_relat_1(X11,X10) | ~v1_relat_1(X11) [subsumption resolution 229,92] {a:1,w:49,wCS:19600,sine:1,thAx:0,allAx:2,thDist:-2}
242. k3_pboole(X10,X9,X11) = X9 | ~r2_pboole(X10,X9,k3_pboole(X10,X9,X11)) | ~v4_relat_1(k3_pboole(X10,X9,X11),X10) | ~v1_relat_1(k3_pboole(X10,X9,X11)) | ~v1_partfun1(X9,X10) | ~v1_funct_1(X9) | ~v4_relat_1(X9,X10) | ~v1_relat_1(X9) | ~v1_partfun1(X11,X10) | ~v1_funct_1(X11) | ~v4_relat_1(X11,X10) | ~v1_relat_1(X11) [subsumption resolution 241,91] {a:1,w:44,wCS:17600,sine:1,thAx:0,allAx:2,thDist:-2}
243. k3_pboole(X10,X9,X11) = X9 | ~r2_pboole(X10,X9,k3_pboole(X10,X9,X11)) | ~v1_relat_1(k3_pboole(X10,X9,X11)) | ~v1_partfun1(X9,X10) | ~v1_funct_1(X9) | ~v4_relat_1(X9,X10) | ~v1_relat_1(X9) | ~v1_partfun1(X11,X10) | ~v1_funct_1(X11) | ~v4_relat_1(X11,X10) | ~v1_relat_1(X11) [subsumption resolution 242,90] {a:1,w:38,wCS:15200,sine:1,thAx:0,allAx:2,thDist:-2}
244. ~r2_pboole(X10,X9,k3_pboole(X10,X9,X11)) | k3_pboole(X10,X9,X11) = X9 | ~v1_partfun1(X9,X10) | ~v1_funct_1(X9) | ~v4_relat_1(X9,X10) | ~v1_relat_1(X9) | ~v1_partfun1(X11,X10) | ~v1_funct_1(X11) | ~v4_relat_1(X11,X10) | ~v1_relat_1(X11) [subsumption resolution 243,89] {a:1,w:33,wCS:13200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
279. k3_pboole(X0,X1,X2) = X1 | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~r2_pboole(X0,X1,X2) | ~r2_pboole(X0,X1,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [resolution 244,74] {a:2,w:64,wCS:25600,sine:1,thAx:0,allAx:3,thDist:-3}
285. k3_pboole(X0,X1,X2) = X1 | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~r2_pboole(X0,X1,X2) | ~r2_pboole(X0,X1,X1) [duplicate literal removal 279] {a:2,w:34,wCS:13600,sine:1,thAx:0,allAx:3,thDist:-3}
286. k3_pboole(X0,X1,X2) = X1 | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~r2_pboole(X0,X1,X2) [subsumption resolution 285,80] {a:2,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
330. ~r6_pboole(sK0,sK1,sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)) <- (~1) [superposition 103,286] {a:3,w:43,wCS:4300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
354. ~r6_pboole(sK0,sK1,sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)) <- (~1, 6) [subsumption resolution 330,128] {a:3,w:40,wCS:4000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
355. ~r6_pboole(sK0,sK1,sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)) <- (~1, 6, 7) [subsumption resolution 354,133] {a:3,w:38,wCS:3800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
356. ~r6_pboole(sK0,sK1,sK1) | ~v1_relat_1(sK1) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)) <- (~1, 6, 7, 8) [subsumption resolution 355,138] {a:3,w:35,wCS:3500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
357. ~r6_pboole(sK0,sK1,sK1) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)) <- (~1, 6, 7, 8, 9) [subsumption resolution 356,143] {a:3,w:33,wCS:3300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
359. 12 <=> r2_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)) [avatar definition]
361. ~r2_pboole(sK0,sK1,k2_pboole(sK0,sK1,sK2)) <- (~12) [avatar component clause 359] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
363. 13 <=> v1_relat_1(k2_pboole(sK0,sK1,sK2)) [avatar definition]
365. ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (~13) [avatar component clause 363] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
367. 14 <=> v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) [avatar definition]
369. ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) <- (~14) [avatar component clause 367] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
371. 15 <=> v1_funct_1(k2_pboole(sK0,sK1,sK2)) [avatar definition]
373. ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) <- (~15) [avatar component clause 371] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
375. 16 <=> v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) [avatar definition]
377. ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) <- (~16) [avatar component clause 375] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
379. 17 <=> r6_pboole(sK0,sK1,sK1) [avatar definition]
381. ~r6_pboole(sK0,sK1,sK1) <- (~17) [avatar component clause 379] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
382. ~12 | ~13 | ~14 | ~15 | ~16 | ~17 | 1 | ~6 | ~7 | ~8 | ~9 [avatar split clause 357,141,136,131,126,101,379,375,371,367,363,359]
428. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~13) [resolution 365,85] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
430. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~13) [subsumption resolution 428,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
431. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~13) [subsumption resolution 430,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
432. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~13) [subsumption resolution 431,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
433. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~13) [subsumption resolution 432,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
434. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~13) [subsumption resolution 433,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
435. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~13) [subsumption resolution 434,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
436. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~13) [subsumption resolution 435,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
437. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~13) [subsumption resolution 436,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
438. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 13 [avatar contradiction clause 437]
460. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~14) [resolution 369,86] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
462. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~14) [subsumption resolution 460,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
463. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~14) [subsumption resolution 462,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
464. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~14) [subsumption resolution 463,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
465. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~14) [subsumption resolution 464,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
466. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~14) [subsumption resolution 465,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
467. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~14) [subsumption resolution 466,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
468. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~14) [subsumption resolution 467,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
469. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~14) [subsumption resolution 468,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
470. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 14 [avatar contradiction clause 469]
472. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~15) [resolution 373,87] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
474. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~15) [subsumption resolution 472,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
475. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~15) [subsumption resolution 474,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
476. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~15) [subsumption resolution 475,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
477. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~15) [subsumption resolution 476,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
478. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~15) [subsumption resolution 477,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
479. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~15) [subsumption resolution 478,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
480. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~15) [subsumption resolution 479,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
481. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~15) [subsumption resolution 480,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
482. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 15 [avatar contradiction clause 481]
725. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~16) [resolution 377,88] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
727. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~16) [subsumption resolution 725,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
728. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~16) [subsumption resolution 727,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
729. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~16) [subsumption resolution 728,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
730. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~16) [subsumption resolution 729,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
731. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~16) [subsumption resolution 730,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
732. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~16) [subsumption resolution 731,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
733. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~16) [subsumption resolution 732,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
734. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~16) [subsumption resolution 733,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
735. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 16 [avatar contradiction clause 734]
737. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~12) [resolution 361,71] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
739. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~12) [subsumption resolution 737,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
740. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~12) [subsumption resolution 739,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
741. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~12) [subsumption resolution 740,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
742. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~12) [subsumption resolution 741,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
743. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~12) [subsumption resolution 742,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
744. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~12) [subsumption resolution 743,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
745. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~12) [subsumption resolution 744,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
746. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~12) [subsumption resolution 745,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
747. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 12 [avatar contradiction clause 746]
749. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~17) [resolution 381,99] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
751. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, ~17) [subsumption resolution 749,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
752. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, ~17) [subsumption resolution 751,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
753. ~v1_relat_1(sK1) <- (6, 7, 8, ~17) [subsumption resolution 752,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
754. $false <- (6, 7, 8, 9, ~17) [subsumption resolution 753,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
755. ~6 | ~7 | ~8 | ~9 | 17 [avatar contradiction clause 754]
756. $false [avatar sat refutation 104,109,114,119,124,129,134,139,144,382,438,470,482,735,747,755] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:325,thDist:-325}
% SZS output end Proof for pboole__t30_pboole
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 603
% Active clauses: 57
% Passive clauses: 112
% Final active clauses: 52
% Final passive clauses: 46
% Main loop iterations started: 66
% 
% >>> Simplifying Inferences
% Duplicate literals: 1284
% Fw subsumption resolutions: 177
% 
% >>> Deletion Inferences
% Simple tautologies: 33
% Equational tautologies: 3
% Forward subsumptions: 88
% 
% >>> Generating Inferences
% Binary resolution: 24
% Unit resulting resolution: 6
% Forward superposition: 57
% Backward superposition: 103
% Self superposition: 4
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 16
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 29
% SAT solver unit clauses: 11
% 
% Memory used [KB]: 5628
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.003 s
% unit resulting resolution: 0.002 s
% superposition: 0.001 s
% other: 0.009 s ( own 0.001 s ) 

