i: [1,167,0,3,1,1,0,0] l11_robbins4
a: [3,250,0,3,1,167]
s: 250
i: [1,168,0,2,1,0,0,1] spc1_boole
a: [3,255,0,2,1,168]
s: 255
i: [1,169,0,2,1,0,0,1] spc2_boole
a: [3,260,0,2,1,169]
s: 260
i: [1,170,0,2,1,0,0,255] fc1_xboole_0
a: [3,265,0,2,1,170]
s: 265
i: [1,178,0,2,1,0,0,255] fc6_ordinal1
a: [3,300,0,2,1,178]
s: 300
i: [1,177,0,2,1,0,0,255] fc6_ordinal1
a: [3,305,0,2,1,177]
s: 305
i: [1,234,0,2,1,0,0,2] rc2_xboole_0
a: [3,347,0,2,1,234]
s: 347
i: [1,235,0,2,1,0,0,3] rc4_ordinal1
a: [3,352,0,2,1,235]
s: 352
i: [1,236,0,2,1,0,0,2] rc1_ordinal1
a: [3,357,0,2,1,236]
s: 357
i: [1,237,0,2,1,0,0,2] rc1_xboole_0
a: [3,362,0,2,1,237]
s: 362
i: [1,241,0,2,1,0,0,255] rc2_ordinal1
a: [3,367,0,2,1,241]
s: 367
i: [1,240,0,2,1,0,0,255] rc2_ordinal1
a: [3,372,0,2,1,240]
s: 372
i: [1,239,0,2,1,0,0,255] rc2_ordinal1
a: [3,377,0,2,1,239]
s: 377
i: [1,238,0,2,1,0,0,255] rc2_ordinal1
a: [3,382,0,2,1,238]
s: 382
i: [1,243,0,2,1,0,0,3] rc5_ordinal1
a: [3,387,0,2,1,243]
s: 387
i: [1,242,0,2,1,0,0,3] rc5_ordinal1
a: [3,392,0,2,1,242]
s: 392
i: [1,245,0,2,1,0,0,255] rc1_relat_1
a: [3,397,0,2,1,245]
s: 397
i: [1,244,0,2,1,0,0,255] rc1_relat_1
a: [3,402,0,2,1,244]
s: 402
i: [1,246,0,2,1,0,0,255] rc2_relat_1
a: [3,407,0,2,1,246]
s: 407
i: [1,185,0,3,1,0,0,2] fc1_subset_1
s: 185
i: [1,221,0,3,1,0,0,2] rc2_subset_1
s: 221
i: [1,224,0,3,1,0,0,1] reflexivity_r1_tarski
s: 224
i: [1,171,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,270,0,3,1,171]
s: 270
i: [1,172,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r2
a: [3,275,0,3,1,172]
s: 275
i: [1,173,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r2_r2
a: [3,280,0,3,1,173]
s: 280
i: [1,174,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,285,0,3,1,174]
s: 285
i: [1,175,0,3,1,0,0,255] t1_numerals
a: [3,290,0,3,1,175]
s: 290
i: [1,181,0,3,1,0,0,255] spc2_numerals
a: [3,310,0,3,1,181]
s: 310
i: [1,180,0,3,1,0,0,255] spc2_numerals
a: [3,315,0,3,1,180]
s: 315
i: [1,184,0,3,1,0,0,255] spc1_numerals
a: [3,325,0,3,1,184]
s: 325
i: [1,183,0,3,1,0,0,255] spc1_numerals
a: [3,330,0,3,1,183]
s: 330
d: [2,409,0,3,1,2,22,315,270]
a: [3,413,0,3,1,409]
s: 413
d: [2,415,0,3,1,2,22,330,270]
a: [3,419,0,3,1,415]
s: 419
i: [1,193,0,4,2,0,0,3] cc6_ordinal1
s: 193
i: [1,197,0,4,2,0,0,255] cc1_ordinal1
s: 197
i: [1,198,0,4,2,0,0,255] cc1_ordinal1
s: 198
d: [2,421,1,2,1,1,40,193,352]
a: [3,426,1,2,1,421]
s: 426
i: [1,200,0,4,2,0,0,255] cc9_ordinal1
s: 200
d: [2,422,1,2,1,1,40,193,387]
a: [3,431,1,2,1,422]
s: 431
i: [1,201,0,4,2,0,0,3] cc7_ordinal1
s: 201
i: [1,202,0,4,2,0,0,255] cc1_relat_1
s: 202
d: [2,433,1,2,1,1,40,197,300]
a: [3,439,1,2,1,433]
s: 439
i: [1,203,0,4,2,0,0,2] cc3_ordinal1
s: 203
d: [2,434,1,2,1,1,40,197,357]
a: [3,444,1,2,1,434]
s: 444
i: [1,219,0,4,1,0,0,2] existence_m1_subset_1
s: 219
i: [1,223,0,4,1,0,0,255] rc3_subset_1
s: 223
d: [2,446,1,2,1,1,40,198,300]
a: [3,452,1,2,1,446]
s: 452
i: [1,176,0,4,1,0,0,255] dt_k5_numbers
a: [3,295,0,4,1,176]
s: 295
d: [2,447,1,2,1,1,40,198,357]
a: [3,457,1,2,1,447]
s: 457
i: [1,179,0,4,1,0,0,255] spc2_numerals
a: [3,320,0,4,1,179]
s: 320
i: [1,182,0,4,1,0,0,255] spc1_numerals
a: [3,335,0,4,1,182]
s: 335
d: [2,471,1,2,1,1,40,200,265]
a: [3,477,1,2,1,471]
s: 477
d: [2,535,0,4,1,2,22,295,270]
a: [3,539,0,4,1,535]
s: 539
d: [2,472,1,2,1,1,40,200,362]
a: [3,482,1,2,1,472]
s: 482
d: [2,541,0,4,1,2,22,320,270]
a: [3,545,0,4,1,541]
s: 545
d: [2,547,0,4,1,2,22,335,270]
a: [3,551,0,4,1,547]
s: 551
d: [2,496,1,2,1,1,40,201,265]
a: [3,502,1,2,1,496]
s: 502
i: [1,189,0,5,2,0,0,2] rc1_subset_1
s: 189
d: [2,497,1,2,1,1,40,201,362]
a: [3,507,1,2,1,497]
s: 507
i: [1,204,0,5,2,0,0,255] t6_boole
s: 204
i: [1,207,0,5,2,0,0,255] cc8_ordinal1
s: 207
d: [2,509,1,2,1,1,40,202,265]
a: [3,515,1,2,1,509]
s: 515
i: [1,216,0,5,2,0,0,255] t7_boole
s: 216
d: [2,510,1,2,1,1,40,202,362]
a: [3,520,1,2,1,510]
s: 520
i: [1,220,0,5,1,0,0,2] rc2_subset_1
s: 220
i: [1,222,0,5,1,0,0,255] rc3_subset_1
s: 222
d: [2,522,1,2,1,1,40,203,265]
a: [3,528,1,2,1,522]
s: 528
i: [1,230,0,7,3,0,0,2] reflexivity_r1_ordinal1
a: [3,342,0,5,2,230]
s: 342
d: [2,523,1,2,1,1,40,203,362]
a: [3,533,1,2,1,523]
s: 533
i: [1,191,0,6,2,0,0,255] rc4_subset_1
s: 191
i: [1,208,0,6,3,0,0,255] cc2_ordinal1
s: 208
d: [2,562,1,2,1,1,40,207,419]
a: [3,569,1,2,1,562]
s: 569
i: [1,211,0,6,2,0,0,255] t1_subset
s: 211
d: [2,563,1,2,1,1,40,207,413]
a: [3,574,1,2,1,563]
s: 574
i: [1,225,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 225
i: [1,188,0,7,2,0,0,2] rc1_subset_1
s: 188
d: [2,459,2,2,1,1,40,426,198]
a: [3,464,2,2,1,459]
s: 464
i: [1,190,0,7,2,0,0,255] rc4_subset_1
s: 190
d: [2,460,2,2,1,1,40,426,197]
a: [3,469,2,2,1,460]
s: 469
i: [1,199,0,7,3,0,0,2] cc5_ordinal1
s: 199
i: [1,213,0,7,2,0,0,1] t3_subset
s: 213
d: [2,484,2,2,1,1,40,431,198]
a: [3,489,2,2,1,484]
s: 489
i: [1,214,0,7,2,0,0,1] t3_subset
s: 214
d: [2,485,2,2,1,1,40,431,197]
a: [3,494,2,2,1,485]
s: 494
i: [1,215,0,7,3,0,0,255] t8_boole
s: 215
i: [1,192,0,8,3,0,0,255] cc10_ordinal1
s: 192
d: [2,659,1,4,2,1,40,192,539]
a: [3,668,1,2,1,659]
s: 668
i: [1,196,0,8,3,0,0,255] cc2_relat_1
s: 196
d: [2,675,1,4,2,1,40,196,539]
a: [3,684,1,2,1,675]
s: 684
i: [1,205,0,8,3,0,0,2] cc1_subset_1
s: 205
i: [1,212,0,8,3,0,0,255] t2_subset
s: 212
d: [2,691,1,4,2,1,40,205,539]
d: [2,697,1,2,1,2,21,691,305]
a: [3,701,1,2,1,697]
s: 701
i: [1,206,0,9,3,0,0,255] cc4_subset_1
s: 206
d: [2,582,2,2,1,1,40,528,198]
a: [3,588,2,2,1,582]
s: 588
i: [1,218,0,9,3,0,0,255] t5_subset
s: 218
i: [1,194,0,10,4,0,0,2] t39_nat_1
s: 194
d: [2,583,2,2,1,1,40,528,197]
a: [3,593,2,2,1,583]
s: 593
i: [1,195,0,10,4,0,0,2] t39_nat_1
s: 195
d: [2,595,2,2,1,1,40,533,198]
a: [3,600,2,2,1,595]
s: 600
i: [1,217,0,10,3,0,0,255] t4_subset
s: 217
i: [1,231,0,10,4,0,0,2] connectedness_r1_ordinal1
s: 231
d: [2,596,2,2,1,1,40,533,197]
a: [3,605,2,2,1,596]
s: 605
i: [1,232,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 232
d: [2,610,2,2,1,1,40,569,193]
a: [3,614,2,2,1,610]
s: 614
i: [1,233,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 233
i: [1,186,0,11,4,0,0,255] cc3_subset_1
s: 186
d: [2,616,2,2,1,1,40,574,193]
a: [3,620,2,2,1,616]
s: 620
i: [1,187,0,11,4,0,0,255] cc2_subset_1
s: 187
d: [2,815,3,2,1,1,40,614,198]
a: [3,820,3,2,1,815]
s: 820
i: [1,229,0,14,4,0,0,255] existence_m2_subset_1
s: 229
d: [2,853,0,12,3,0,21,229,205]
s: 853
d: [2,816,3,2,1,1,40,614,197]
a: [3,825,3,2,1,816]
s: 825
i: [1,226,0,15,5,0,0,255] dt_m2_subset_1
s: 226
d: [2,840,3,2,1,1,40,620,198]
a: [3,845,3,2,1,840]
s: 845
d: [2,854,0,13,4,0,21,226,205]
s: 854
i: [1,227,0,15,5,0,0,255] redefinition_m2_subset_1
s: 227
d: [2,841,3,2,1,1,40,620,197]
a: [3,850,3,2,1,841]
s: 850
d: [2,860,0,13,4,0,21,227,205]
s: 860
d: [2,473,1,3,1,0,40,200,221]
s: 473
i: [1,228,0,15,5,0,0,255] redefinition_m2_subset_1
s: 228
d: [2,865,0,13,4,0,21,228,205]
s: 865
d: [2,498,1,3,1,0,40,201,221]
s: 498
d: [2,511,1,3,1,0,40,202,221]
s: 511
d: [2,524,1,3,1,0,40,203,221]
s: 524
d: [2,565,1,3,1,0,40,207,219]
a: [3,579,1,3,1,565]
s: 579
d: [2,644,1,3,1,1,40,213,539]
a: [3,653,1,3,1,644]
s: 653
d: [2,623,1,4,2,1,40,199,325]
a: [3,637,1,2,1,623]
s: 637
d: [2,704,1,5,2,1,40,212,325]
a: [3,719,1,3,1,704]
s: 719
d: [2,705,1,5,2,1,40,212,419]
d: [2,721,1,3,1,2,21,705,305]
a: [3,725,1,3,1,721]
s: 725
d: [2,706,1,5,2,1,40,212,310]
a: [3,730,1,3,1,706]
s: 730
d: [2,707,1,5,2,1,40,212,413]
d: [2,732,1,3,1,2,21,707,305]
a: [3,736,1,3,1,732]
s: 736
d: [2,709,1,5,2,1,40,212,290]
d: [2,744,1,3,1,2,21,709,305]
a: [3,748,1,3,1,744]
s: 748
d: [2,770,1,7,3,1,40,194,275]
d: [2,775,1,5,2,2,21,770,574]
d: [2,776,1,3,1,3,21,775,569]
a: [3,780,1,3,1,776]
s: 780
d: [2,932,3,0,0,4,52,614,620,250,780,232]
e: 932
d: [2,933,2,7,3,1,40,780,232]
d: [2,936,2,4,2,2,21,933,250]
d: [2,937,2,2,1,3,21,936,620]
d: [2,938,2,0,0,4,21,937,614]
e: 938
d: [2,624,1,4,2,1,40,199,419]
d: [2,639,1,2,1,2,21,624,300]
d: [2,626,1,4,2,1,40,199,413]
d: [2,642,1,2,1,2,21,626,300]
f: 938,776,642,639,563,562,415,409,183,180,178,172,171,167
% Refutation found. Thanks to Tanya!
% SZS status Theorem for robbins4__l11_robbins4
% SZS output start Proof for robbins4__l11_robbins4
1. r1_tarski(1,2) [input l11_robbins4]
2. ~r1_tarski(1,2) [negated conjecture 1]
13. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
15. v3_ordinal1(k4_ordinal1) & ~v1_xboole_0(k4_ordinal1) [input fc6_ordinal1]
27. k4_ordinal1 = k5_numbers [input redefinition_k5_numbers]
41. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (m1_subset_1(X1,X0) => v3_ordinal1(X1))) [input cc5_ordinal1]
59. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
65. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
66. m1_subset_1(2,k1_numbers) & m1_subset_1(2,k5_numbers) & m2_subset_1(2,k1_numbers,k5_numbers) & v2_xxreal_0(2) [input spc2_numerals]
67. ! [X0] : (v7_ordinal1(X0) => ! [X1] : (v7_ordinal1(X1) => (r1_xxreal_0(X0,X1) <=> r1_ordinal1(X0,X1)))) [input t39_nat_1]
69. r1_xxreal_0(1,2) [input rqLessOrEqual__r1_xxreal_0__r1_r2]
71. ~r1_tarski(1,2) [flattening 2]
80. m1_subset_1(2,k1_numbers) & m1_subset_1(2,k5_numbers) & m2_subset_1(2,k1_numbers,k5_numbers) [pure predicate removal 66]
81. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) [pure predicate removal 65]
96. ! [X0] : (! [X1] : ((r1_xxreal_0(X0,X1) <=> r1_ordinal1(X0,X1)) | ~v7_ordinal1(X1)) | ~v7_ordinal1(X0)) [ennf transformation 67]
99. ! [X0] : (! [X1] : (v3_ordinal1(X1) | ~m1_subset_1(X1,X0)) | ~v3_ordinal1(X0)) [ennf transformation 41]
107. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 13]
133. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 59]
134. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 133]
139. ! [X0] : (! [X1] : (((r1_xxreal_0(X0,X1) | ~r1_ordinal1(X0,X1)) & (r1_ordinal1(X0,X1) | ~r1_xxreal_0(X0,X1))) | ~v7_ordinal1(X1)) | ~v7_ordinal1(X0)) [nnf transformation 96]
150. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 134]
167. ~r1_tarski(1,2) [cnf transformation 71] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
171. k4_ordinal1 = k5_numbers [cnf transformation 27] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
172. r1_xxreal_0(1,2) [cnf transformation 69] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
178. v3_ordinal1(k4_ordinal1) [cnf transformation 15] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
180. m1_subset_1(2,k5_numbers) [cnf transformation 80] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
183. m1_subset_1(1,k5_numbers) [cnf transformation 81] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
194. ~r1_xxreal_0(X0,X1) | r1_ordinal1(X0,X1) | ~v7_ordinal1(X1) | ~v7_ordinal1(X0) [cnf transformation 139] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
199. ~m1_subset_1(X1,X0) | v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 99] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
207. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 107] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
232. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 150] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
248. 1 <=> r1_tarski(1,2) [avatar definition]
250. ~r1_tarski(1,2) <- (~1) [avatar component clause 248] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
251. ~1 [avatar split clause 167,248]
268. 5 <=> k4_ordinal1 = k5_numbers [avatar definition]
270. k4_ordinal1 = k5_numbers <- (5) [avatar component clause 268] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
271. 5 [avatar split clause 171,268]
273. 6 <=> r1_xxreal_0(1,2) [avatar definition]
275. r1_xxreal_0(1,2) <- (6) [avatar component clause 273] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
276. 6 [avatar split clause 172,273]
298. 11 <=> v3_ordinal1(k4_ordinal1) [avatar definition]
300. v3_ordinal1(k4_ordinal1) <- (11) [avatar component clause 298] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
301. 11 [avatar split clause 178,298]
313. 14 <=> m1_subset_1(2,k5_numbers) [avatar definition]
315. m1_subset_1(2,k5_numbers) <- (14) [avatar component clause 313] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
316. 14 [avatar split clause 180,313]
328. 17 <=> m1_subset_1(1,k5_numbers) [avatar definition]
330. m1_subset_1(1,k5_numbers) <- (17) [avatar component clause 328] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
331. 17 [avatar split clause 183,328]
409. m1_subset_1(2,k4_ordinal1) <- (5, 14) [forward demodulation 315,270] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
411. 34 <=> m1_subset_1(2,k4_ordinal1) [avatar definition]
413. m1_subset_1(2,k4_ordinal1) <- (34) [avatar component clause 411] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
414. 34 | ~5 | ~14 [avatar split clause 409,313,268,411]
415. m1_subset_1(1,k4_ordinal1) <- (5, 17) [forward demodulation 330,270] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
417. 35 <=> m1_subset_1(1,k4_ordinal1) [avatar definition]
419. m1_subset_1(1,k4_ordinal1) <- (35) [avatar component clause 417] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
420. 35 | ~5 | ~17 [avatar split clause 415,328,268,417]
562. v7_ordinal1(1) <- (35) [resolution 207,419] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
563. v7_ordinal1(2) <- (34) [resolution 207,413] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
567. 57 <=> v7_ordinal1(1) [avatar definition]
569. v7_ordinal1(1) <- (57) [avatar component clause 567] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
570. 57 | ~35 [avatar split clause 562,417,567]
572. 58 <=> v7_ordinal1(2) [avatar definition]
574. v7_ordinal1(2) <- (58) [avatar component clause 572] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
575. 58 | ~34 [avatar split clause 563,411,572]
612. 64 <=> v3_ordinal1(1) [avatar definition]
614. v3_ordinal1(1) <- (64) [avatar component clause 612] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
618. 65 <=> v3_ordinal1(2) [avatar definition]
620. v3_ordinal1(2) <- (65) [avatar component clause 618] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
624. v3_ordinal1(1) | ~v3_ordinal1(k4_ordinal1) <- (35) [resolution 199,419] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:2,thDist:-2}
626. v3_ordinal1(2) | ~v3_ordinal1(k4_ordinal1) <- (34) [resolution 199,413] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:2,thDist:-2}
639. v3_ordinal1(1) <- (11, 35) [subsumption resolution 624,300] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
640. 64 | ~11 | ~35 [avatar split clause 639,417,298,612]
642. v3_ordinal1(2) <- (11, 34) [subsumption resolution 626,300] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
643. 65 | ~11 | ~34 [avatar split clause 642,411,298,618]
770. r1_ordinal1(1,2) | ~v7_ordinal1(2) | ~v7_ordinal1(1) <- (6) [resolution 194,275] {a:1,w:7,wCS:2800,sine:2,thAx:0,allAx:2,thDist:-2}
775. r1_ordinal1(1,2) | ~v7_ordinal1(1) <- (6, 58) [subsumption resolution 770,574] {a:1,w:5,wCS:2000,sine:2,thAx:0,allAx:2,thDist:-2}
776. r1_ordinal1(1,2) <- (6, 57, 58) [subsumption resolution 775,569] {a:1,w:3,wCS:1200,sine:2,thAx:0,allAx:2,thDist:-2}
778. 79 <=> r1_ordinal1(1,2) [avatar definition]
780. r1_ordinal1(1,2) <- (79) [avatar component clause 778] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
781. 79 | ~6 | ~57 | ~58 [avatar split clause 776,572,567,273,778]
933. r1_tarski(1,2) | ~v3_ordinal1(2) | ~v3_ordinal1(1) <- (79) [resolution 780,232] {a:2,w:7,wCS:2800,sine:1,thAx:0,allAx:3,thDist:-3}
936. ~v3_ordinal1(2) | ~v3_ordinal1(1) <- (~1, 79) [subsumption resolution 933,250] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
937. ~v3_ordinal1(1) <- (~1, 65, 79) [subsumption resolution 936,620] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
938. $false <- (~1, 64, 65, 79) [subsumption resolution 937,614] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
939. 1 | ~64 | ~65 | ~79 [avatar contradiction clause 938]
940. $false [avatar sat refutation 251,271,276,301,316,331,414,420,570,575,640,643,781,939] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:400,thDist:-400}
% SZS output end Proof for robbins4__l11_robbins4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 69
% >>> Preprocessing
% Introduced skolems: 14
% Pure predicates: 5
% 
% >>> Saturation
% Initial clauses: 80
% Generated clauses: 399
% Active clauses: 123
% Passive clauses: 211
% Final active clauses: 123
% Final passive clauses: 78
% Main loop iterations started: 133
% 
% >>> Simplifying Inferences
% Duplicate literals: 15
% Fw subsumption resolutions: 31
% Fw demodulations: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Forward subsumptions: 45
% 
% >>> Generating Inferences
% Binary resolution: 172
% Unit resulting resolution: 1
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 16
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 98
% SAT solver unit clauses: 32
% SAT solver binary clauses: 41
% 
% Memory used [KB]: 9978
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% SAT solver time: 0.001 s
% unit resulting resolution: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.006 s ( own 0.001 s ) 

