i: [1,121,0,2,1,1,0,0] t59_partfun1
a: [3,187,0,2,1,121]
s: 187
i: [1,120,0,2,1,1,0,0] t59_partfun1
a: [3,192,0,2,1,120]
s: 192
i: [1,122,0,3,1,1,0,0] t59_partfun1
a: [3,182,0,3,1,122]
s: 182
i: [1,123,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,197,0,2,1,123]
s: 197
i: [1,124,0,2,1,0,0,1] fc1_xboole_0
a: [3,202,0,2,1,124]
s: 202
i: [1,159,0,2,1,0,0,2] rc2_xboole_0
a: [3,219,0,2,1,159]
s: 219
i: [1,160,0,2,1,0,0,2] rc1_xboole_0
a: [3,224,0,2,1,160]
s: 224
i: [1,162,0,2,1,0,0,255] rc7_funct_1
a: [3,229,0,2,1,162]
s: 229
i: [1,161,0,2,1,0,0,255] rc7_funct_1
a: [3,234,0,2,1,161]
s: 234
i: [1,164,0,2,1,0,0,1] rc1_relat_1
a: [3,239,0,2,1,164]
s: 239
i: [1,163,0,2,1,0,0,1] rc1_relat_1
a: [3,244,0,2,1,163]
s: 244
i: [1,167,0,2,1,0,0,255] rc4_funct_1
a: [3,249,0,2,1,167]
s: 249
i: [1,166,0,2,1,0,0,255] rc4_funct_1
a: [3,254,0,2,1,166]
s: 254
i: [1,165,0,2,1,0,0,255] rc4_funct_1
a: [3,259,0,2,1,165]
s: 259
i: [1,168,0,2,1,0,0,255] rc2_relat_1
a: [3,264,0,2,1,168]
s: 264
i: [1,170,0,2,1,0,0,255] rc3_funct_1
a: [3,269,0,2,1,170]
s: 269
i: [1,169,0,2,1,0,0,255] rc3_funct_1
a: [3,274,0,2,1,169]
s: 274
i: [1,172,0,2,1,0,0,1] rc1_funct_1
a: [3,279,0,2,1,172]
s: 279
i: [1,171,0,2,1,0,0,1] rc1_funct_1
a: [3,284,0,2,1,171]
s: 284
i: [1,175,0,2,1,0,0,255] rc5_funct_1
a: [3,289,0,2,1,175]
s: 289
i: [1,174,0,2,1,0,0,255] rc5_funct_1
a: [3,294,0,2,1,174]
s: 294
i: [1,173,0,2,1,0,0,255] rc5_funct_1
a: [3,299,0,2,1,173]
s: 299
i: [1,177,0,2,1,0,0,255] rc2_funct_1
a: [3,304,0,2,1,177]
s: 304
i: [1,176,0,2,1,0,0,255] rc2_funct_1
a: [3,309,0,2,1,176]
s: 309
i: [1,125,0,3,1,0,0,255] d2_xboole_0
a: [3,207,0,3,1,125]
s: 207
r: 197
i: [1,129,0,4,2,0,0,255] cc7_funct_1
s: 129
d: [2,312,1,2,1,1,40,129,224]
a: [3,317,1,2,1,312]
s: 317
i: [1,130,0,4,2,0,0,1] cc1_funct_1
s: 130
d: [2,313,1,2,1,1,40,129,202]
a: [3,322,1,2,1,313]
s: 322
i: [1,131,0,4,2,0,0,1] cc1_relat_1
s: 131
i: [1,150,0,4,1,0,0,255] existence_m1_subset_1
s: 150
d: [2,324,1,2,1,1,40,130,224]
a: [3,329,1,2,1,324]
s: 329
i: [1,126,0,5,1,0,0,1] t1_boole
s: 126
d: [2,325,1,2,1,1,40,130,202]
a: [3,334,1,2,1,325]
s: 334
i: [1,132,0,5,2,0,0,1] t6_boole
s: 132
i: [1,149,0,5,2,0,0,255] t7_boole
s: 149
d: [2,353,1,5,2,1,44,182,132]
s: 353
i: [1,151,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 151
d: [2,336,1,2,1,1,40,131,224]
a: [3,341,1,2,1,336]
s: 341
i: [1,146,0,6,2,0,0,255] t1_subset
s: 146
i: [1,153,0,6,2,0,0,2] fc4_xboole_0
s: 153
d: [2,337,1,2,1,1,40,131,202]
a: [3,346,1,2,1,337]
s: 346
i: [1,154,0,6,2,0,0,2] fc5_xboole_0
s: 154
i: [1,155,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 155
i: [1,127,0,7,3,0,0,255] cc8_funct_1
s: 127
i: [1,128,0,7,3,0,0,255] cc8_funct_1
s: 128
d: [2,359,1,5,2,0,40,127,150]
s: 359
i: [1,148,0,7,3,0,0,2] t8_boole
s: 148
d: [2,360,1,5,2,0,40,128,150]
s: 360
i: [1,152,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 152
i: [1,156,0,11,5,0,0,1] reflexivity_r1_partfun1
a: [3,214,0,7,3,156]
s: 214
d: [2,369,2,6,3,2,40,214,353]
d: [2,370,2,4,2,3,21,369,192]
d: [2,371,2,2,1,4,21,370,187]
a: [3,375,2,2,1,371]
s: 375
i: [1,137,0,8,4,0,0,255] cc4_funct_1
s: 137
d: [2,377,0,6,3,0,21,137,130]
d: [2,378,0,4,2,0,21,377,131]
s: 378
i: [1,147,0,8,3,0,0,255] t2_subset
s: 147
i: [1,158,0,8,3,0,0,2] fc3_relat_1
s: 158
d: [2,379,1,2,1,1,40,378,224]
a: [3,384,1,2,1,379]
s: 384
i: [1,157,0,14,6,0,0,1] symmetry_r1_partfun1
s: 157
d: [2,380,1,2,1,1,40,378,202]
a: [3,389,1,2,1,380]
s: 389
i: [1,142,0,15,6,0,0,1] t51_partfun1
s: 142
i: [1,143,0,15,6,0,0,1] t51_partfun1
s: 143
d: [2,361,1,5,1,0,44,152,126]
s: 361
i: [1,144,0,18,6,0,0,1] t51_partfun1
s: 144
d: [2,362,1,5,1,0,44,152,126]
s: 362
i: [1,178,0,19,7,0,0,1] t51_partfun1
s: 178
d: [2,412,0,15,6,0,21,178,158]
s: 412
d: [2,363,1,5,1,0,44,126,152]
s: 363
d: [2,366,1,5,1,0,44,126,152]
s: 366
d: [2,391,1,6,2,0,40,147,150]
s: 391
d: [2,348,1,7,3,0,44,132,132]
s: 348
d: [2,350,1,7,2,0,44,126,132]
s: 350
d: [2,425,2,6,2,0,40,391,155]
s: 425
d: [2,413,1,13,6,0,44,412,126]
d: [2,420,1,11,5,0,18,413]
d: [2,421,1,9,4,1,21,420,334]
d: [2,422,1,7,3,2,21,421,346]
s: 422
d: [2,400,2,7,2,0,44,361,132]
s: 400
d: [2,408,1,26,11,0,44,143,144]
d: [2,411,1,15,6,0,18,408]
s: 411
d: [2,415,1,15,6,0,44,412,152]
s: 415
d: [2,417,2,13,6,0,44,412,361]
d: [2,418,2,11,5,0,18,417]
d: [2,423,2,9,4,1,21,418,334]
d: [2,424,2,7,3,2,21,423,346]
s: 424
d: [2,491,3,0,0,5,52,192,187,182,424]
e: 491
d: [2,492,3,4,2,3,40,424,182]
d: [2,498,3,2,1,4,21,492,187]
d: [2,499,3,0,0,5,21,498,192]
e: 499
d: [2,493,3,6,3,3,40,424,353]
d: [2,501,3,4,2,4,21,493,187]
d: [2,502,3,2,1,5,21,501,192]
d: [2,503,3,0,0,6,21,502,202]
e: 503
f: 503,337,325,124,120,121,122
% Refutation found. Thanks to Tanya!
% SZS status Theorem for partfun1__t59_partfun1
% SZS output start Proof for partfun1__t59_partfun1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => r1_partfun1(k1_xboole_0,X0)) [input t59_partfun1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => r1_partfun1(k1_xboole_0,X0)) [negated conjecture 1]
4. ! [X0] : (v1_xboole_0(X0) => v1_funct_1(X0)) [input cc1_funct_1]
5. ! [X0] : (v1_xboole_0(X0) => v1_relat_1(X0)) [input cc1_relat_1]
14. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
21. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
22. ! [X0,X1] : ((v1_relat_1(X1) & v1_relat_1(X0)) => v1_relat_1(k2_xboole_0(X0,X1))) [input fc3_relat_1]
38. ! [X0] : k2_xboole_0(X0,k1_xboole_0) = X0 [input t1_boole]
41. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) <=> ? [X2] : (k2_xboole_0(X0,X1) = X2 & v1_funct_1(X2) & v1_relat_1(X2))))) [input t51_partfun1]
42. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
59. ? [X0] : (~r1_partfun1(k1_xboole_0,X0) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
60. ? [X0] : (~r1_partfun1(k1_xboole_0,X0) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 59]
63. ! [X0] : (v1_funct_1(X0) | ~v1_xboole_0(X0)) [ennf transformation 4]
64. ! [X0] : (v1_relat_1(X0) | ~v1_xboole_0(X0)) [ennf transformation 5]
65. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 42]
76. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> ? [X2] : (k2_xboole_0(X0,X1) = X2 & v1_funct_1(X2) & v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 41]
77. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> ? [X2] : (k2_xboole_0(X0,X1) = X2 & v1_funct_1(X2) & v1_relat_1(X2))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 76]
90. ! [X0,X1] : (v1_relat_1(k2_xboole_0(X0,X1)) | (~v1_relat_1(X1) | ~v1_relat_1(X0))) [ennf transformation 22]
91. ! [X0,X1] : (v1_relat_1(k2_xboole_0(X0,X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0)) [flattening 90]
92. ? [X0] : (~r1_partfun1(k1_xboole_0,X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (~r1_partfun1(k1_xboole_0,sK0) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
93. ~r1_partfun1(k1_xboole_0,sK0) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 60,92]
94. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ! [X2] : (k2_xboole_0(X0,X1) != X2 | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X2] : (k2_xboole_0(X0,X1) = X2 & v1_funct_1(X2) & v1_relat_1(X2)) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 77]
95. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ! [X2] : (k2_xboole_0(X0,X1) != X2 | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X3] : (k2_xboole_0(X0,X1) = X3 & v1_funct_1(X3) & v1_relat_1(X3)) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [rectify 94]
96. ! [X1,X0] : (? [X3] : (k2_xboole_0(X0,X1) = X3 & v1_funct_1(X3) & v1_relat_1(X3)) => (k2_xboole_0(X0,X1) = sK1(X0,X1) & v1_funct_1(sK1(X0,X1)) & v1_relat_1(sK1(X0,X1)))) [choice axiom]
97. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ! [X2] : (k2_xboole_0(X0,X1) != X2 | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & ((k2_xboole_0(X0,X1) = sK1(X0,X1) & v1_funct_1(sK1(X0,X1)) & v1_relat_1(sK1(X0,X1))) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [skolemisation 95,96]
120. v1_relat_1(sK0) [cnf transformation 93] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
121. v1_funct_1(sK0) [cnf transformation 93] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
122. ~r1_partfun1(k1_xboole_0,sK0) [cnf transformation 93] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
124. v1_xboole_0(k1_xboole_0) [cnf transformation 21] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
126. k2_xboole_0(X0,k1_xboole_0) = X0 [cnf transformation 38] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
130. ~v1_xboole_0(X0) | v1_funct_1(X0) [cnf transformation 63] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
131. ~v1_xboole_0(X0) | v1_relat_1(X0) [cnf transformation 64] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
132. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 65] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
145. r1_partfun1(X0,X1) | k2_xboole_0(X0,X1) != X2 | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 97] {a:0,w:20,wCS:8000,sine:1,thAx:0,allAx:2,thDist:-2}
152. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 14] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
158. v1_relat_1(k2_xboole_0(X0,X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 91] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
178. r1_partfun1(X0,X1) | ~v1_funct_1(k2_xboole_0(X0,X1)) | ~v1_relat_1(k2_xboole_0(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [equality resolution 145] {a:0,w:19,wCS:7600,sine:1,thAx:0,allAx:2,thDist:-2}
180. 1 <=> r1_partfun1(k1_xboole_0,sK0) [avatar definition]
182. ~r1_partfun1(k1_xboole_0,sK0) <- (~1) [avatar component clause 180] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
183. ~1 [avatar split clause 122,180]
185. 2 <=> v1_funct_1(sK0) [avatar definition]
187. v1_funct_1(sK0) <- (2) [avatar component clause 185] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
188. 2 [avatar split clause 121,185]
190. 3 <=> v1_relat_1(sK0) [avatar definition]
192. v1_relat_1(sK0) <- (3) [avatar component clause 190] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
193. 3 [avatar split clause 120,190]
200. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
202. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 200] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
203. 5 [avatar split clause 124,200]
325. v1_funct_1(k1_xboole_0) <- (5) [resolution 130,202] {a:1,w:2,wCS:800,sine:1,thAx:0,allAx:2,thDist:-2}
332. 31 <=> v1_funct_1(k1_xboole_0) [avatar definition]
334. v1_funct_1(k1_xboole_0) <- (31) [avatar component clause 332] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
335. 31 | ~5 [avatar split clause 325,200,332]
337. v1_relat_1(k1_xboole_0) <- (5) [resolution 131,202] {a:1,w:2,wCS:800,sine:1,thAx:0,allAx:2,thDist:-2}
344. 33 <=> v1_relat_1(k1_xboole_0) [avatar definition]
346. v1_relat_1(k1_xboole_0) <- (33) [avatar component clause 344] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
347. 33 | ~5 [avatar split clause 337,200,344]
353. ~r1_partfun1(X0,sK0) | ~v1_xboole_0(X0) <- (~1) [superposition 182,132] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
361. k2_xboole_0(k1_xboole_0,X0) = X0 [superposition 152,126] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
412. ~v1_funct_1(k2_xboole_0(X0,X1)) | r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 178,158] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
417. ~v1_funct_1(X6) | r1_partfun1(k1_xboole_0,X6) | ~v1_funct_1(X6) | ~v1_relat_1(X6) | ~v1_funct_1(k1_xboole_0) | ~v1_relat_1(k1_xboole_0) [superposition 412,361] {a:2,w:13,wCS:5200,sine:1,thAx:0,allAx:4,thDist:-4}
418. ~v1_funct_1(X6) | r1_partfun1(k1_xboole_0,X6) | ~v1_relat_1(X6) | ~v1_funct_1(k1_xboole_0) | ~v1_relat_1(k1_xboole_0) [duplicate literal removal 417] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:4,thDist:-4}
423. ~v1_funct_1(X6) | r1_partfun1(k1_xboole_0,X6) | ~v1_relat_1(X6) | ~v1_relat_1(k1_xboole_0) <- (31) [subsumption resolution 418,334] {a:2,w:9,wCS:3600,sine:1,thAx:0,allAx:4,thDist:-4}
424. r1_partfun1(k1_xboole_0,X6) | ~v1_funct_1(X6) | ~v1_relat_1(X6) <- (31, 33) [subsumption resolution 423,346] {a:2,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
493. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_xboole_0(k1_xboole_0) <- (~1, 31, 33) [resolution 424,353] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
501. ~v1_relat_1(sK0) | ~v1_xboole_0(k1_xboole_0) <- (~1, 2, 31, 33) [subsumption resolution 493,187] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
502. ~v1_xboole_0(k1_xboole_0) <- (~1, 2, 3, 31, 33) [subsumption resolution 501,192] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
503. $false <- (~1, 2, 3, 5, 31, 33) [subsumption resolution 502,202] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
504. 1 | ~2 | ~3 | ~5 | ~31 | ~33 [avatar contradiction clause 503]
505. $false [avatar sat refutation 183,188,193,203,335,347,504] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:143,thDist:-143}
% SZS output end Proof for partfun1__t59_partfun1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 43
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 58
% Generated clauses: 239
% Active clauses: 71
% Passive clauses: 91
% Final active clauses: 70
% Final passive clauses: 14
% Main loop iterations started: 77
% 
% >>> Simplifying Inferences
% Duplicate literals: 43
% Fw subsumption resolutions: 26
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 29
% Forward subsumptions: 41
% 
% >>> Generating Inferences
% Binary resolution: 22
% Unit resulting resolution: 1
% Forward superposition: 43
% Backward superposition: 33
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 40
% SAT solver unit clauses: 25
% SAT solver binary clauses: 9
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

