i: [1,54,0,2,1,1,0,0] t116_pboole
a: [3,101,0,2,1,54]
s: 101
i: [1,52,0,2,1,1,0,0] t116_pboole
a: [3,111,0,2,1,52]
s: 111
i: [1,50,0,2,1,1,0,0] t116_pboole
a: [3,121,0,2,1,50]
s: 121
i: [1,48,0,2,1,1,0,0] t116_pboole
a: [3,131,0,2,1,48]
s: 131
i: [1,55,0,3,1,1,0,0] t116_pboole
a: [3,96,0,3,1,55]
s: 96
i: [1,53,0,3,1,1,0,0] t116_pboole
a: [3,106,0,3,1,53]
s: 106
i: [1,51,0,3,1,1,0,0] t116_pboole
a: [3,116,0,3,1,51]
s: 116
i: [1,49,0,3,1,1,0,0] t116_pboole
a: [3,126,0,3,1,49]
s: 126
i: [1,56,0,4,1,1,0,0] t116_pboole
a: [3,91,0,4,1,56]
s: 91
i: [1,80,0,2,1,0,0,1] rc1_funct_1
a: [3,136,0,2,1,80]
s: 136
i: [1,79,0,2,1,0,0,1] rc1_funct_1
a: [3,141,0,2,1,79]
s: 141
i: [1,57,0,10,1,1,0,0] t116_pboole
a: [3,86,0,10,1,57]
s: 86
i: [1,60,0,3,1,0,0,1] rc1_pboole
s: 60
i: [1,62,0,3,1,0,0,1] rc1_pboole
s: 62
i: [1,61,0,4,1,0,0,1] rc1_pboole
s: 61
i: [1,63,0,4,1,0,0,1] rc1_pboole
s: 63
i: [1,81,0,24,9,0,0,255] redefinition_r6_pboole
d: [2,82,0,14,5,0,18,81]
s: 82
i: [1,64,0,24,9,0,0,1] reflexivity_r6_pboole
s: 64
i: [1,67,0,25,9,0,0,1] dt_k4_pboole
s: 67
i: [1,69,0,25,9,0,0,1] dt_k4_pboole
s: 69
i: [1,71,0,25,9,0,0,1] dt_k2_pboole
s: 71
i: [1,73,0,25,9,0,0,1] dt_k2_pboole
s: 73
i: [1,65,0,26,9,0,0,255] idempotence_k2_pboole
s: 65
i: [1,68,0,26,9,0,0,1] dt_k4_pboole
s: 68
i: [1,70,0,26,9,0,0,1] dt_k4_pboole
s: 70
i: [1,72,0,26,9,0,0,1] dt_k2_pboole
s: 72
i: [1,74,0,26,9,0,0,1] dt_k2_pboole
s: 74
i: [1,77,0,27,10,0,0,255] redefinition_r6_pboole
s: 77
i: [1,75,0,28,10,0,0,1] symmetry_r5_pboole
s: 75
i: [1,76,0,28,10,0,0,1] symmetry_r6_pboole
s: 76
i: [1,66,0,29,9,0,0,255] commutativity_k2_pboole
s: 66
d: [2,153,1,24,9,1,40,75,91]
d: [2,154,1,21,8,2,21,153,96]
d: [2,155,1,19,7,3,21,154,101]
d: [2,156,1,16,6,4,21,155,106]
d: [2,157,1,14,5,5,21,156,111]
d: [2,158,1,11,4,6,21,157,116]
d: [2,159,1,9,3,7,21,158,121]
d: [2,160,1,6,2,8,21,159,126]
d: [2,161,1,4,1,9,21,160,131]
a: [3,165,1,4,1,161]
s: 165
i: [1,59,0,31,10,0,0,1] t115_pboole
s: 59
d: [2,187,1,62,18,0,40,59,77]
d: [2,188,1,52,14,0,18,187]
d: [2,194,1,46,13,0,21,188,70]
d: [2,195,1,41,12,0,21,194,69]
d: [2,196,1,35,11,0,21,195,68]
d: [2,197,1,30,10,0,21,196,67]
s: 197
i: [1,58,0,33,9,0,0,1] t75_pboole
s: 58
d: [2,240,1,50,17,0,44,58,65]
d: [2,249,1,30,9,0,18,240]
s: 249
d: [2,198,2,46,10,1,44,86,197]
d: [2,209,2,43,9,2,21,198,96]
d: [2,210,2,41,8,3,21,209,101]
d: [2,211,2,38,7,4,21,210,106]
d: [2,212,2,36,6,5,21,211,111]
a: [3,216,2,7,1,212]
s: 216
d: [2,186,1,63,18,0,40,59,76]
d: [2,189,1,53,14,0,18,186]
d: [2,190,1,47,13,0,21,189,70]
d: [2,191,1,42,12,0,21,190,69]
d: [2,192,1,36,11,0,21,191,68]
d: [2,193,1,31,10,0,21,192,67]
s: 193
d: [2,242,1,53,17,0,44,58,66]
d: [2,247,1,33,9,0,18,242]
s: 247
d: [2,241,1,53,17,0,44,58,66]
d: [2,248,1,33,9,0,18,241]
s: 248
d: [2,239,1,88,17,0,40,58,77]
d: [2,254,1,82,16,0,21,239,70]
d: [2,255,1,77,15,0,21,254,69]
d: [2,256,1,71,14,0,21,255,68]
d: [2,257,1,66,13,0,21,256,67]
s: 257
d: [2,301,2,61,13,1,44,86,257]
d: [2,331,2,59,12,2,21,301,101]
d: [2,332,2,56,11,3,21,331,106]
d: [2,333,2,54,10,4,21,332,111]
d: [2,334,2,51,9,5,21,333,116]
d: [2,335,2,49,8,6,21,334,121]
d: [2,336,2,46,7,7,21,335,126]
d: [2,337,2,44,6,8,21,336,131]
d: [2,338,2,41,5,9,21,337,96]
a: [3,342,2,8,1,338]
s: 342
d: [2,238,1,89,17,0,40,58,76]
d: [2,250,1,83,16,0,21,238,70]
d: [2,251,1,78,15,0,21,250,69]
d: [2,252,1,72,14,0,21,251,68]
d: [2,253,1,67,13,0,21,252,67]
s: 253
a: [3,220,2,5,1,212]
s: 220
d: [2,415,3,0,0,9,52,131,111,121,101,126,116,106,96,220,71]
e: 415
d: [2,416,3,20,8,1,40,220,71]
d: [2,418,3,17,7,2,21,416,96]
d: [2,419,3,15,6,3,21,418,101]
d: [2,420,3,12,5,4,21,419,106]
d: [2,421,3,10,4,5,21,420,111]
d: [2,422,3,7,3,6,21,421,116]
d: [2,423,3,5,2,7,21,422,121]
d: [2,424,3,2,1,8,21,423,126]
d: [2,425,3,0,0,9,21,424,131]
e: 425
r: 220
a: [3,219,2,5,1,212]
s: 219
a: [3,224,2,6,1,212]
s: 224
d: [2,427,3,0,0,9,52,131,111,121,101,126,116,106,96,224,72]
e: 427
d: [2,428,3,20,8,1,40,224,72]
d: [2,430,3,17,7,2,21,428,96]
d: [2,431,3,15,6,3,21,430,101]
d: [2,432,3,12,5,4,21,431,106]
d: [2,433,3,10,4,5,21,432,111]
d: [2,434,3,7,3,6,21,433,116]
d: [2,435,3,5,2,7,21,434,121]
d: [2,436,3,2,1,8,21,435,126]
d: [2,437,3,0,0,9,21,436,131]
e: 437
r: 224
a: [3,228,2,5,1,212]
s: 228
d: [2,439,3,0,0,9,52,131,111,121,101,126,116,106,96,228,73]
e: 439
d: [2,440,3,20,8,1,40,228,73]
d: [2,442,3,17,7,2,21,440,96]
d: [2,443,3,15,6,3,21,442,101]
d: [2,444,3,12,5,4,21,443,106]
d: [2,445,3,10,4,5,21,444,111]
d: [2,446,3,7,3,6,21,445,116]
d: [2,447,3,5,2,7,21,446,121]
d: [2,448,3,2,1,8,21,447,126]
d: [2,449,3,0,0,9,21,448,131]
e: 449
r: 228
a: [3,227,2,5,1,212]
s: 227
a: [3,223,2,6,1,212]
s: 223
a: [3,232,2,6,1,212]
s: 232
d: [2,451,3,0,0,9,52,131,111,121,101,126,116,106,96,232,74]
e: 451
d: [2,452,3,20,8,1,40,232,74]
d: [2,454,3,17,7,2,21,452,96]
d: [2,455,3,15,6,3,21,454,101]
d: [2,456,3,12,5,4,21,455,106]
d: [2,457,3,10,4,5,21,456,111]
d: [2,458,3,7,3,6,21,457,116]
d: [2,459,3,5,2,7,21,458,121]
d: [2,460,3,2,1,8,21,459,126]
d: [2,461,3,0,0,9,21,460,131]
e: 461
r: 232
r: 342
a: [3,231,2,6,1,212]
s: 231
a: [3,341,2,8,1,338]
s: 341
a: [3,346,2,9,1,338]
s: 346
d: [2,465,3,0,0,9,52,111,101,106,96,219,227,223,231,346,68]
e: 465
d: [2,466,3,32,8,1,40,346,68]
d: [2,471,3,29,7,2,21,466,96]
d: [2,472,3,27,6,3,21,471,101]
d: [2,473,3,24,5,4,21,472,106]
d: [2,474,3,22,4,5,21,473,111]
d: [2,475,3,16,3,6,21,474,231]
d: [2,476,3,11,2,7,21,475,227]
d: [2,477,3,5,1,8,21,476,223]
d: [2,478,3,0,0,9,21,477,219]
e: 478
r: 346
a: [3,350,2,8,1,338]
s: 350
d: [2,480,3,0,0,9,52,111,101,106,96,219,227,223,231,350,69]
e: 480
d: [2,481,3,32,8,1,40,350,69]
d: [2,486,3,29,7,2,21,481,96]
d: [2,487,3,27,6,3,21,486,101]
d: [2,488,3,24,5,4,21,487,106]
d: [2,489,3,22,4,5,21,488,111]
d: [2,490,3,16,3,6,21,489,231]
d: [2,491,3,11,2,7,21,490,227]
d: [2,492,3,5,1,8,21,491,223]
d: [2,493,3,0,0,9,21,492,219]
e: 493
r: 350
a: [3,349,2,8,1,338]
s: 349
a: [3,345,2,9,1,338]
s: 345
a: [3,354,2,9,1,338]
s: 354
d: [2,499,3,0,0,9,52,111,101,106,96,219,227,223,231,354,70]
e: 499
d: [2,500,3,32,8,1,40,354,70]
d: [2,505,3,29,7,2,21,500,96]
d: [2,506,3,27,6,3,21,505,101]
d: [2,507,3,24,5,4,21,506,106]
d: [2,508,3,22,4,5,21,507,111]
d: [2,509,3,16,3,6,21,508,231]
d: [2,510,3,11,2,7,21,509,227]
d: [2,511,3,5,1,8,21,510,223]
d: [2,512,3,0,0,9,21,511,219]
e: 512
r: 354
a: [3,358,2,7,1,338]
s: 358
d: [2,514,3,0,0,10,52,131,111,121,101,126,116,106,96,91,358,59]
e: 514
d: [2,515,3,24,9,1,40,358,59]
d: [2,518,3,20,8,2,21,515,91]
d: [2,519,3,17,7,3,21,518,96]
d: [2,520,3,15,6,4,21,519,101]
d: [2,521,3,12,5,5,21,520,106]
d: [2,522,3,10,4,6,21,521,111]
d: [2,523,3,7,3,7,21,522,116]
d: [2,524,3,5,2,8,21,523,121]
d: [2,525,3,2,1,9,21,524,126]
d: [2,526,3,0,0,10,21,525,131]
e: 526
d: [2,380,3,32,8,1,40,342,67]
d: [2,384,3,29,7,2,21,380,96]
d: [2,385,3,27,6,3,21,384,101]
d: [2,386,3,24,5,4,21,385,106]
d: [2,387,3,22,4,5,21,386,111]
f: 526,512,493,478,461,449,437,425,387,338,48,49,50,51,52,53,54,55,56,57
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pboole__t116_pboole
% SZS output start Proof for pboole__t116_pboole
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r5_pboole(X0,X1,X2) => r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),X1)))) [input t116_pboole]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r5_pboole(X0,X1,X2) => r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),X1)))) [negated conjecture 1]
4. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2)))) [input dt_k2_pboole]
5. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k4_pboole(X0,X1,X2),X0) & v1_funct_1(k4_pboole(X0,X1,X2)) & v4_relat_1(k4_pboole(X0,X1,X2),X0) & v1_relat_1(k4_pboole(X0,X1,X2)))) [input dt_k4_pboole]
9. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (r6_pboole(X0,X1,X2) <=> X1 = X2)) [input redefinition_r6_pboole]
13. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r5_pboole(X0,X1,X2) => r6_pboole(X0,k4_pboole(X0,X1,X2),X1)))) [input t115_pboole]
14. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),k4_pboole(X0,X1,X2)))) [input t75_pboole]
15. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r5_pboole(X0,X1,X2) => r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),X1)))) [flattening 2]
16. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),k4_pboole(X0,X1,X2)))) [flattening 14]
17. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r5_pboole(X0,X1,X2) => r6_pboole(X0,k4_pboole(X0,X1,X2),X1)))) [flattening 13]
18. ? [X0,X1] : (? [X2] : ((~r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),X1) & r5_pboole(X0,X1,X2)) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 15]
19. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),X1) & r5_pboole(X0,X1,X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 18]
20. ! [X0,X1] : (! [X2] : (r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),k4_pboole(X0,X1,X2)) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 16]
21. ! [X0,X1] : (! [X2] : (r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),k4_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 20]
22. ! [X0,X1] : (! [X2] : ((r6_pboole(X0,k4_pboole(X0,X1,X2),X1) | ~r5_pboole(X0,X1,X2)) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 17]
23. ! [X0,X1] : (! [X2] : (r6_pboole(X0,k4_pboole(X0,X1,X2),X1) | ~r5_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 22]
30. ! [X0,X1,X2] : ((v1_partfun1(k4_pboole(X0,X1,X2),X0) & v1_funct_1(k4_pboole(X0,X1,X2)) & v4_relat_1(k4_pboole(X0,X1,X2),X0) & v1_relat_1(k4_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 5]
31. ! [X0,X1,X2] : ((v1_partfun1(k4_pboole(X0,X1,X2),X0) & v1_funct_1(k4_pboole(X0,X1,X2)) & v4_relat_1(k4_pboole(X0,X1,X2),X0) & v1_relat_1(k4_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 30]
32. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 4]
33. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 32]
38. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 9]
39. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 38]
40. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),X1) & r5_pboole(X0,X1,X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (~r6_pboole(sK0,k4_pboole(sK0,k2_pboole(sK0,sK1,X2),X2),sK1) & r5_pboole(sK0,sK1,X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
41. ? [X2] : (~r6_pboole(sK0,k4_pboole(sK0,k2_pboole(sK0,sK1,X2),X2),sK1) & r5_pboole(sK0,sK1,X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (~r6_pboole(sK0,k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK1) & r5_pboole(sK0,sK1,sK2) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
42. (~r6_pboole(sK0,k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK1) & r5_pboole(sK0,sK1,sK2) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 19,41,40]
45. ! [X0,X1,X2] : (((r6_pboole(X0,X1,X2) | X1 != X2) & (X1 = X2 | ~r6_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 39]
48. v1_relat_1(sK1) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. v4_relat_1(sK1,sK0) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. v1_funct_1(sK1) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. v1_partfun1(sK1,sK0) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. v1_relat_1(sK2) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
53. v4_relat_1(sK2,sK0) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. v1_funct_1(sK2) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. v1_partfun1(sK2,sK0) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. r5_pboole(sK0,sK1,sK2) [cnf transformation 42] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
57. ~r6_pboole(sK0,k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK1) [cnf transformation 42] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. r6_pboole(X0,k4_pboole(X0,k2_pboole(X0,X1,X2),X2),k4_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 21] {a:0,w:33,wCS:13200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. r6_pboole(X0,k4_pboole(X0,X1,X2),X1) | ~r5_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 23] {a:0,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. v1_relat_1(k4_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 31] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. v4_relat_1(k4_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 31] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. v1_funct_1(k4_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 31] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. v1_partfun1(k4_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 31] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. v1_relat_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 33] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. v4_relat_1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 33] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. v1_funct_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 33] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
74. v1_partfun1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 33] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
77. ~r6_pboole(X0,X1,X2) | X1 = X2 | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 45] {a:0,w:27,wCS:10800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
84. 1 <=> r6_pboole(sK0,k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK1) [avatar definition]
86. ~r6_pboole(sK0,k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK1) <- (~1) [avatar component clause 84] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
87. ~1 [avatar split clause 57,84]
89. 2 <=> r5_pboole(sK0,sK1,sK2) [avatar definition]
91. r5_pboole(sK0,sK1,sK2) <- (2) [avatar component clause 89] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. 2 [avatar split clause 56,89]
94. 3 <=> v1_partfun1(sK2,sK0) [avatar definition]
96. v1_partfun1(sK2,sK0) <- (3) [avatar component clause 94] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
97. 3 [avatar split clause 55,94]
99. 4 <=> v1_funct_1(sK2) [avatar definition]
101. v1_funct_1(sK2) <- (4) [avatar component clause 99] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
102. 4 [avatar split clause 54,99]
104. 5 <=> v4_relat_1(sK2,sK0) [avatar definition]
106. v4_relat_1(sK2,sK0) <- (5) [avatar component clause 104] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
107. 5 [avatar split clause 53,104]
109. 6 <=> v1_relat_1(sK2) [avatar definition]
111. v1_relat_1(sK2) <- (6) [avatar component clause 109] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
112. 6 [avatar split clause 52,109]
114. 7 <=> v1_partfun1(sK1,sK0) [avatar definition]
116. v1_partfun1(sK1,sK0) <- (7) [avatar component clause 114] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
117. 7 [avatar split clause 51,114]
119. 8 <=> v1_funct_1(sK1) [avatar definition]
121. v1_funct_1(sK1) <- (8) [avatar component clause 119] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
122. 8 [avatar split clause 50,119]
124. 9 <=> v4_relat_1(sK1,sK0) [avatar definition]
126. v4_relat_1(sK1,sK0) <- (9) [avatar component clause 124] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
127. 9 [avatar split clause 49,124]
129. 10 <=> v1_relat_1(sK1) [avatar definition]
131. v1_relat_1(sK1) <- (10) [avatar component clause 129] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
132. 10 [avatar split clause 48,129]
218. 15 <=> v1_relat_1(k2_pboole(sK0,sK1,sK2)) [avatar definition]
219. v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (15) [avatar component clause 218] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
220. ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (~15) [avatar component clause 218] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
222. 16 <=> v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) [avatar definition]
223. v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) <- (16) [avatar component clause 222] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
224. ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) <- (~16) [avatar component clause 222] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
226. 17 <=> v1_funct_1(k2_pboole(sK0,sK1,sK2)) [avatar definition]
227. v1_funct_1(k2_pboole(sK0,sK1,sK2)) <- (17) [avatar component clause 226] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
228. ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) <- (~17) [avatar component clause 226] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
230. 18 <=> v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) [avatar definition]
231. v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) <- (18) [avatar component clause 230] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
232. ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) <- (~18) [avatar component clause 230] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
239. ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | k4_pboole(X4,k2_pboole(X4,X5,X3),X3) = k4_pboole(X4,X5,X3) | ~v1_partfun1(k4_pboole(X4,X5,X3),X4) | ~v1_funct_1(k4_pboole(X4,X5,X3)) | ~v4_relat_1(k4_pboole(X4,X5,X3),X4) | ~v1_relat_1(k4_pboole(X4,X5,X3)) | ~v1_partfun1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_funct_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) | ~v4_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) [resolution 58,77] {a:1,w:88,wCS:35200,sine:1,thAx:0,allAx:2,thDist:-2}
254. ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | k4_pboole(X4,k2_pboole(X4,X5,X3),X3) = k4_pboole(X4,X5,X3) | ~v1_funct_1(k4_pboole(X4,X5,X3)) | ~v4_relat_1(k4_pboole(X4,X5,X3),X4) | ~v1_relat_1(k4_pboole(X4,X5,X3)) | ~v1_partfun1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_funct_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) | ~v4_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) [subsumption resolution 239,70] {a:1,w:82,wCS:32800,sine:1,thAx:0,allAx:2,thDist:-2}
255. ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | k4_pboole(X4,k2_pboole(X4,X5,X3),X3) = k4_pboole(X4,X5,X3) | ~v4_relat_1(k4_pboole(X4,X5,X3),X4) | ~v1_relat_1(k4_pboole(X4,X5,X3)) | ~v1_partfun1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_funct_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) | ~v4_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) [subsumption resolution 254,69] {a:1,w:77,wCS:30800,sine:1,thAx:0,allAx:2,thDist:-2}
256. ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | k4_pboole(X4,k2_pboole(X4,X5,X3),X3) = k4_pboole(X4,X5,X3) | ~v1_relat_1(k4_pboole(X4,X5,X3)) | ~v1_partfun1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_funct_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) | ~v4_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) [subsumption resolution 255,68] {a:1,w:71,wCS:28400,sine:1,thAx:0,allAx:2,thDist:-2}
257. k4_pboole(X4,k2_pboole(X4,X5,X3),X3) = k4_pboole(X4,X5,X3) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~v1_partfun1(X3,X4) | ~v1_partfun1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_funct_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) | ~v4_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3),X4) | ~v1_relat_1(k4_pboole(X4,k2_pboole(X4,X5,X3),X3)) [subsumption resolution 256,67] {a:1,w:66,wCS:26400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
301. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) | ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~1) [superposition 86,257] {a:2,w:61,wCS:6100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
331. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) | ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~1, 4) [subsumption resolution 301,101] {a:2,w:59,wCS:5900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
332. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) | ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~1, 4, 5) [subsumption resolution 331,106] {a:2,w:56,wCS:5600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
333. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) | ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~1, 4, 5, 6) [subsumption resolution 332,111] {a:2,w:54,wCS:5400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
334. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) | ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~1, 4, 5, 6, 7) [subsumption resolution 333,116] {a:2,w:51,wCS:5100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
335. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) | ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~1, 4, 5, 6, 7, 8) [subsumption resolution 334,121] {a:2,w:49,wCS:4900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
336. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) | ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~1, 4, 5, 6, 7, 8, 9) [subsumption resolution 335,126] {a:2,w:46,wCS:4600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
337. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) | ~v1_partfun1(sK2,sK0) | ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) | ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~1, 4, 5, 6, 7, 8, 9, 10) [subsumption resolution 336,131] {a:2,w:44,wCS:4400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
338. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) | ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) | ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) | ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~1, 3, 4, 5, 6, 7, 8, 9, 10) [subsumption resolution 337,96] {a:2,w:41,wCS:4100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
340. 20 <=> v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) [avatar definition]
342. ~v1_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~20) [avatar component clause 340] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
344. 21 <=> v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) [avatar definition]
346. ~v4_relat_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) <- (~21) [avatar component clause 344] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
348. 22 <=> v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) [avatar definition]
350. ~v1_funct_1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2)) <- (~22) [avatar component clause 348] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
352. 23 <=> v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) [avatar definition]
354. ~v1_partfun1(k4_pboole(sK0,k2_pboole(sK0,sK1,sK2),sK2),sK0) <- (~23) [avatar component clause 352] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
356. 24 <=> r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) [avatar definition]
358. ~r6_pboole(sK0,k4_pboole(sK0,sK1,sK2),sK1) <- (~24) [avatar component clause 356] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
359. ~20 | ~21 | ~22 | ~23 | ~24 | 1 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 [avatar split clause 338,129,124,119,114,109,104,99,94,84,356,352,348,344,340]
380. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (~20) [resolution 342,67] {a:3,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
384. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, ~20) [subsumption resolution 380,96] {a:3,w:29,wCS:2900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
385. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, ~20) [subsumption resolution 384,101] {a:3,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
386. ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, ~20) [subsumption resolution 385,106] {a:3,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
387. ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, ~20) [subsumption resolution 386,111] {a:3,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
388. ~15 | ~16 | ~17 | ~18 | ~3 | ~4 | ~5 | ~6 | 20 [avatar split clause 387,340,109,104,99,94,230,226,222,218]
416. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~15) [resolution 220,71] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
418. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, ~15) [subsumption resolution 416,96] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
419. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, ~15) [subsumption resolution 418,101] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
420. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, ~15) [subsumption resolution 419,106] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
421. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, ~15) [subsumption resolution 420,111] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
422. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, ~15) [subsumption resolution 421,116] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
423. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, 8, ~15) [subsumption resolution 422,121] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
424. ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, 8, 9, ~15) [subsumption resolution 423,126] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
425. $false <- (3, 4, 5, 6, 7, 8, 9, 10, ~15) [subsumption resolution 424,131] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
426. ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | 15 [avatar contradiction clause 425]
428. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~16) [resolution 224,72] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
430. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, ~16) [subsumption resolution 428,96] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
431. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, ~16) [subsumption resolution 430,101] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
432. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, ~16) [subsumption resolution 431,106] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
433. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, ~16) [subsumption resolution 432,111] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
434. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, ~16) [subsumption resolution 433,116] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
435. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, 8, ~16) [subsumption resolution 434,121] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
436. ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, 8, 9, ~16) [subsumption resolution 435,126] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
437. $false <- (3, 4, 5, 6, 7, 8, 9, 10, ~16) [subsumption resolution 436,131] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
438. ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | 16 [avatar contradiction clause 437]
440. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~17) [resolution 228,73] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
442. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, ~17) [subsumption resolution 440,96] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
443. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, ~17) [subsumption resolution 442,101] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
444. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, ~17) [subsumption resolution 443,106] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
445. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, ~17) [subsumption resolution 444,111] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
446. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, ~17) [subsumption resolution 445,116] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
447. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, 8, ~17) [subsumption resolution 446,121] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
448. ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, 8, 9, ~17) [subsumption resolution 447,126] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
449. $false <- (3, 4, 5, 6, 7, 8, 9, 10, ~17) [subsumption resolution 448,131] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
450. ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | 17 [avatar contradiction clause 449]
452. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~18) [resolution 232,74] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
454. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, ~18) [subsumption resolution 452,96] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
455. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, ~18) [subsumption resolution 454,101] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
456. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, ~18) [subsumption resolution 455,106] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
457. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, ~18) [subsumption resolution 456,111] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
458. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, ~18) [subsumption resolution 457,116] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
459. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, 8, ~18) [subsumption resolution 458,121] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
460. ~v1_relat_1(sK1) <- (3, 4, 5, 6, 7, 8, 9, ~18) [subsumption resolution 459,126] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
461. $false <- (3, 4, 5, 6, 7, 8, 9, 10, ~18) [subsumption resolution 460,131] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
462. ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | 18 [avatar contradiction clause 461]
466. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (~21) [resolution 346,68] {a:3,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
471. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, ~21) [subsumption resolution 466,96] {a:3,w:29,wCS:2900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
472. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, ~21) [subsumption resolution 471,101] {a:3,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
473. ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, ~21) [subsumption resolution 472,106] {a:3,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
474. ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, ~21) [subsumption resolution 473,111] {a:3,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
475. ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, 18, ~21) [subsumption resolution 474,231] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
476. ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, 17, 18, ~21) [subsumption resolution 475,227] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
477. ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, 16, 17, 18, ~21) [subsumption resolution 476,223] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
478. $false <- (3, 4, 5, 6, 15, 16, 17, 18, ~21) [subsumption resolution 477,219] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
479. ~3 | ~4 | ~5 | ~6 | ~15 | ~16 | ~17 | ~18 | 21 [avatar contradiction clause 478]
481. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (~22) [resolution 350,69] {a:3,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
486. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, ~22) [subsumption resolution 481,96] {a:3,w:29,wCS:2900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
487. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, ~22) [subsumption resolution 486,101] {a:3,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
488. ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, ~22) [subsumption resolution 487,106] {a:3,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
489. ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, ~22) [subsumption resolution 488,111] {a:3,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
490. ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, 18, ~22) [subsumption resolution 489,231] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
491. ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, 17, 18, ~22) [subsumption resolution 490,227] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
492. ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, 16, 17, 18, ~22) [subsumption resolution 491,223] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
493. $false <- (3, 4, 5, 6, 15, 16, 17, 18, ~22) [subsumption resolution 492,219] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
494. ~3 | ~4 | ~5 | ~6 | ~15 | ~16 | ~17 | ~18 | 22 [avatar contradiction clause 493]
500. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (~23) [resolution 354,70] {a:3,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
505. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, ~23) [subsumption resolution 500,96] {a:3,w:29,wCS:2900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
506. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, ~23) [subsumption resolution 505,101] {a:3,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
507. ~v1_relat_1(sK2) | ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, ~23) [subsumption resolution 506,106] {a:3,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
508. ~v1_partfun1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, ~23) [subsumption resolution 507,111] {a:3,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
509. ~v1_funct_1(k2_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, 18, ~23) [subsumption resolution 508,231] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
510. ~v4_relat_1(k2_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, 17, 18, ~23) [subsumption resolution 509,227] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
511. ~v1_relat_1(k2_pboole(sK0,sK1,sK2)) <- (3, 4, 5, 6, 16, 17, 18, ~23) [subsumption resolution 510,223] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
512. $false <- (3, 4, 5, 6, 15, 16, 17, 18, ~23) [subsumption resolution 511,219] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
513. ~3 | ~4 | ~5 | ~6 | ~15 | ~16 | ~17 | ~18 | 23 [avatar contradiction clause 512]
515. ~r5_pboole(sK0,sK1,sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~24) [resolution 358,59] {a:3,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
518. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~24) [subsumption resolution 515,91] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
519. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~24) [subsumption resolution 518,96] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
520. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~24) [subsumption resolution 519,101] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
521. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~24) [subsumption resolution 520,106] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
522. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~24) [subsumption resolution 521,111] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
523. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~24) [subsumption resolution 522,116] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
524. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~24) [subsumption resolution 523,121] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
525. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, 9, ~24) [subsumption resolution 524,126] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
526. $false <- (2, 3, 4, 5, 6, 7, 8, 9, 10, ~24) [subsumption resolution 525,131] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
527. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | 24 [avatar contradiction clause 526]
537. $false [avatar sat refutation 87,92,97,102,107,112,117,122,127,132,359,388,426,438,450,462,479,494,513,527] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:425,thDist:-425}
% SZS output end Proof for pboole__t116_pboole
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 33
% Generated clauses: 391
% Active clauses: 56
% Passive clauses: 79
% Final active clauses: 48
% Final passive clauses: 21
% Main loop iterations started: 58
% 
% >>> Simplifying Inferences
% Duplicate literals: 414
% Fw subsumption resolutions: 162
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Equational tautologies: 2
% Forward subsumptions: 41
% 
% >>> Generating Inferences
% Binary resolution: 25
% Unit resulting resolution: 8
% Forward superposition: 39
% Backward superposition: 28
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 15
% Sat splitting refutations: 16
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 34
% SAT solver unit clauses: 12
% 
% Memory used [KB]: 9850
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% forward subsumption: 0.002 s
% forward subsumption resolution: 0.001 s
% other: 0.005 s ( own 0.001 s ) 

