i: [1,74,0,8,2,1,0,0] t13_mcart_1
a: [3,108,0,4,1,74]
s: 108
i: [1,73,0,6,1,1,0,0] t13_mcart_1
a: [3,117,0,6,1,73]
s: 117
i: [1,75,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,122,0,2,1,75]
s: 122
i: [1,76,0,2,1,0,0,255] fc1_xboole_0
a: [3,127,0,2,1,76]
s: 127
i: [1,97,0,2,1,0,0,2] rc2_xboole_0
a: [3,137,0,2,1,97]
s: 137
i: [1,98,0,2,1,0,0,2] rc1_xboole_0
a: [3,142,0,2,1,98]
s: 142
i: [1,100,0,2,1,0,0,2] rc1_relat_1
a: [3,147,0,2,1,100]
s: 147
i: [1,99,0,2,1,0,0,2] rc1_relat_1
a: [3,152,0,2,1,99]
s: 152
i: [1,101,0,2,1,0,0,255] rc2_relat_1
a: [3,157,0,2,1,101]
s: 157
i: [1,78,0,3,1,0,0,1] fc2_xboole_0
s: 78
i: [1,77,0,3,1,0,0,255] d2_xboole_0
a: [3,132,0,3,1,77]
s: 132
r: 122
i: [1,79,0,4,2,0,0,2] cc1_relat_1
s: 79
i: [1,93,0,4,1,0,0,255] existence_m1_subset_1
s: 93
d: [2,160,1,2,1,1,40,79,142]
a: [3,165,1,2,1,160]
s: 165
i: [1,94,0,4,1,0,0,1] fc6_relat_1
s: 94
i: [1,103,0,4,1,0,0,1] d1_tarski
s: 103
d: [2,161,1,2,1,1,40,79,127]
a: [3,170,1,2,1,161]
s: 170
i: [1,80,0,5,2,0,0,255] t6_boole
s: 80
i: [1,90,0,5,2,0,0,1] t7_boole
s: 90
i: [1,84,0,6,2,0,0,255] t1_subset
s: 84
i: [1,95,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 95
d: [2,175,1,5,1,1,40,90,117]
a: [3,180,1,5,1,175]
s: 180
i: [1,89,0,7,3,0,0,1] t8_boole
s: 89
d: [2,182,1,6,1,1,40,84,117]
a: [3,187,1,6,1,182]
s: 187
i: [1,104,0,7,2,0,0,1] d1_tarski
s: 104
i: [1,83,0,8,3,0,0,255] t2_subset
s: 83
d: [2,189,1,6,1,1,40,95,117]
a: [3,194,1,6,1,189]
s: 194
i: [1,96,0,8,3,0,0,1] fc10_subset_1
s: 96
d: [2,183,1,4,1,0,40,84,103]
s: 183
i: [1,91,0,9,2,0,0,1] t10_mcart_1
s: 91
d: [2,200,1,0,0,2,52,108,117,91]
e: 200
d: [2,201,1,4,1,1,40,91,117]
d: [2,203,1,0,0,2,21,201,108]
e: 203
r: 108
a: [3,107,0,4,1,74]
s: 107
d: [2,208,1,2,1,1,40,107,90]
a: [3,222,1,2,1,208]
s: 222
a: [3,112,0,4,1,74]
s: 112
d: [2,206,1,4,1,1,40,107,95]
a: [3,212,1,4,1,206]
s: 212
i: [1,92,0,9,2,0,0,1] t10_mcart_1
s: 92
i: [1,87,0,14,3,0,0,1] d1_tarski
s: 87
d: [2,207,1,4,1,1,40,107,84]
a: [3,217,1,4,1,207]
s: 217
i: [1,88,0,14,3,0,0,1] d1_tarski
s: 88
d: [2,224,1,5,1,1,40,92,117]
a: [3,228,1,5,1,224]
s: 228
d: [2,237,2,0,0,2,52,112,228,104]
e: 237
d: [2,238,2,4,1,1,40,228,104]
d: [2,243,2,0,0,2,21,238,112]
e: 243
f: 243,224,201,73,74
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t13_mcart_1
% SZS output start Proof for mcart_1__t13_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(X1,k1_tarski(X2))) => (k2_xtuple_0(X0) = X2 & r2_hidden(k1_xtuple_0(X0),X1))) [input t13_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(X1,k1_tarski(X2))) => (k2_xtuple_0(X0) = X2 & r2_hidden(k1_xtuple_0(X0),X1))) [negated conjecture 1]
7. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
25. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [input t10_mcart_1]
31. ~! [X0,X1,X2] : (r2_hidden(X0,k2_zfmisc_1(X1,k1_tarski(X2))) => (k2_xtuple_0(X0) = X2 & r2_hidden(k1_xtuple_0(X0),X1))) [flattening 2]
34. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 7]
37. ! [X0,X1,X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [flattening 25]
41. ? [X0,X1,X2] : ((k2_xtuple_0(X0) != X2 | ~r2_hidden(k1_xtuple_0(X0),X1)) & r2_hidden(X0,k2_zfmisc_1(X1,k1_tarski(X2)))) [ennf transformation 31]
53. ! [X0,X1,X2] : ((r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1)) | ~r2_hidden(X0,k2_zfmisc_1(X1,X2))) [ennf transformation 37]
57. ? [X0,X1,X2] : ((k2_xtuple_0(X0) != X2 | ~r2_hidden(k1_xtuple_0(X0),X1)) & r2_hidden(X0,k2_zfmisc_1(X1,k1_tarski(X2)))) => ((sK2 != k2_xtuple_0(sK0) | ~r2_hidden(k1_xtuple_0(sK0),sK1)) & r2_hidden(sK0,k2_zfmisc_1(sK1,k1_tarski(sK2)))) [choice axiom]
58. (sK2 != k2_xtuple_0(sK0) | ~r2_hidden(k1_xtuple_0(sK0),sK1)) & r2_hidden(sK0,k2_zfmisc_1(sK1,k1_tarski(sK2))) [skolemisation 41,57]
59. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 34]
60. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 59]
61. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
62. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 60,61]
73. r2_hidden(sK0,k2_zfmisc_1(sK1,k1_tarski(sK2))) [cnf transformation 58] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
74. sK2 != k2_xtuple_0(sK0) | ~r2_hidden(k1_xtuple_0(sK0),sK1) [cnf transformation 58] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
85. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 62] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
91. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k1_xtuple_0(X0),X1) [cnf transformation 53] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
92. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k2_xtuple_0(X0),X2) [cnf transformation 53] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 85] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
106. 1 <=> r2_hidden(k1_xtuple_0(sK0),sK1) [avatar definition]
110. 2 <=> sK2 = k2_xtuple_0(sK0) [avatar definition]
112. sK2 != k2_xtuple_0(sK0) <- (~2) [avatar component clause 110] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
113. ~1 | ~2 [avatar split clause 74,110,106]
115. 3 <=> r2_hidden(sK0,k2_zfmisc_1(sK1,k1_tarski(sK2))) [avatar definition]
117. r2_hidden(sK0,k2_zfmisc_1(sK1,k1_tarski(sK2))) <- (3) [avatar component clause 115] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
118. 3 [avatar split clause 73,115]
201. r2_hidden(k1_xtuple_0(sK0),sK1) <- (3) [resolution 91,117] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
205. 1 | ~3 [avatar split clause 201,115,106]
224. r2_hidden(k2_xtuple_0(sK0),k1_tarski(sK2)) <- (3) [resolution 92,117] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
226. 20 <=> r2_hidden(k2_xtuple_0(sK0),k1_tarski(sK2)) [avatar definition]
228. r2_hidden(k2_xtuple_0(sK0),k1_tarski(sK2)) <- (20) [avatar component clause 226] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
229. 20 | ~3 [avatar split clause 224,115,226]
238. sK2 = k2_xtuple_0(sK0) <- (20) [resolution 228,104] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
243. $false <- (~2, 20) [subsumption resolution 238,112] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
244. 2 | ~20 [avatar contradiction clause 243]
245. $false [avatar sat refutation 113,118,205,229,244] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:97,thDist:-97}
% SZS output end Proof for mcart_1__t13_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 87
% Active clauses: 39
% Passive clauses: 48
% Final active clauses: 37
% Final passive clauses: 9
% Main loop iterations started: 39
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 28
% Unit resulting resolution: 2
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 9
% SAT solver binary clauses: 15
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

