i: [1,200,0,2,1,1,0,0] t68_xxreal_1
a: [3,320,0,2,1,200]
s: 320
i: [1,199,0,2,1,1,0,0] t68_xxreal_1
a: [3,325,0,2,1,199]
s: 325
i: [1,198,0,2,1,1,0,0] t68_xxreal_1
a: [3,330,0,2,1,198]
s: 330
i: [1,197,0,2,1,1,0,0] t68_xxreal_1
a: [3,335,0,2,1,197]
s: 335
i: [1,203,0,6,2,1,0,0] t68_xxreal_1
a: [3,301,0,3,1,203]
s: 301
i: [1,202,0,3,1,1,0,0] t68_xxreal_1
a: [3,310,0,3,1,202]
s: 310
i: [1,201,0,7,1,1,0,0] t68_xxreal_1
a: [3,315,0,7,1,201]
s: 315
i: [1,204,0,2,1,0,0,255] fc5_numbers
a: [3,340,0,2,1,204]
s: 340
i: [1,205,0,2,1,0,0,255] fc2_membered
a: [3,345,0,2,1,205]
s: 345
i: [1,206,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,350,0,2,1,206]
s: 350
i: [1,207,0,2,1,0,0,255] fc1_xboole_0
a: [3,355,0,2,1,207]
s: 355
i: [1,264,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,364,0,2,1,264]
s: 364
d: [2,498,1,0,0,2,40,364,320]
e: 498
d: [2,497,1,0,0,2,40,364,325]
e: 497
d: [2,496,1,0,0,2,40,364,330]
e: 496
d: [2,495,1,0,0,2,40,364,335]
e: 495
d: [2,494,1,0,0,2,52,335,364]
e: 494
r: 364
i: [1,272,0,2,1,0,0,2] rc2_xboole_0
a: [3,372,0,2,1,272]
s: 372
i: [1,273,0,2,1,0,0,255] rc1_xreal_0
a: [3,377,0,2,1,273]
s: 377
i: [1,274,0,2,1,0,0,2] rc1_xboole_0
a: [3,382,0,2,1,274]
s: 382
i: [1,275,0,2,1,0,0,1] rc1_xxreal_0
a: [3,387,0,2,1,275]
s: 387
i: [1,277,0,2,1,0,0,255] rc1_membered
a: [3,392,0,2,1,277]
s: 392
i: [1,276,0,2,1,0,0,255] rc1_membered
a: [3,397,0,2,1,276]
s: 397
i: [1,279,0,2,1,0,0,255] rc2_membered
a: [3,402,0,2,1,279]
s: 402
i: [1,278,0,2,1,0,0,255] rc2_membered
a: [3,407,0,2,1,278]
s: 407
i: [1,281,0,2,1,0,0,255] rc3_membered
a: [3,412,0,2,1,281]
s: 412
i: [1,280,0,2,1,0,0,255] rc3_membered
a: [3,417,0,2,1,280]
s: 417
i: [1,284,0,2,1,0,0,255] rc3_xreal_0
a: [3,422,0,2,1,284]
s: 422
i: [1,283,0,2,1,0,0,255] rc3_xreal_0
a: [3,427,0,2,1,283]
s: 427
i: [1,282,0,2,1,0,0,255] rc3_xreal_0
a: [3,432,0,2,1,282]
s: 432
i: [1,287,0,2,1,0,0,255] rc2_xreal_0
a: [3,437,0,2,1,287]
s: 437
i: [1,286,0,2,1,0,0,255] rc2_xreal_0
a: [3,442,0,2,1,286]
s: 442
i: [1,285,0,2,1,0,0,255] rc2_xreal_0
a: [3,447,0,2,1,285]
s: 447
i: [1,290,0,2,1,0,0,255] rc4_xreal_0
a: [3,452,0,2,1,290]
s: 452
i: [1,289,0,2,1,0,0,255] rc4_xreal_0
a: [3,457,0,2,1,289]
s: 457
i: [1,288,0,2,1,0,0,255] rc4_xreal_0
a: [3,462,0,2,1,288]
s: 462
i: [1,292,0,2,1,0,0,2] rc4_xxreal_0
a: [3,467,0,2,1,292]
s: 467
i: [1,291,0,2,1,0,0,2] rc4_xxreal_0
a: [3,472,0,2,1,291]
s: 472
i: [1,294,0,2,1,0,0,255] rc2_xxreal_0
a: [3,477,0,2,1,294]
s: 477
i: [1,293,0,2,1,0,0,255] rc2_xxreal_0
a: [3,482,0,2,1,293]
s: 482
i: [1,296,0,2,1,0,0,255] rc3_xxreal_0
a: [3,487,0,2,1,296]
s: 487
i: [1,295,0,2,1,0,0,255] rc3_xxreal_0
a: [3,492,0,2,1,295]
s: 492
i: [1,262,0,3,1,0,0,255] reflexivity_r1_tarski
s: 262
i: [1,208,0,3,1,0,0,255] d2_xboole_0
a: [3,360,0,3,1,208]
s: 360
r: 350
i: [1,209,0,4,2,0,0,255] cc2_xreal_0
s: 209
i: [1,210,0,4,2,0,0,255] cc2_xxreal_0
s: 210
i: [1,212,0,4,2,0,0,255] cc2_membered
s: 212
i: [1,213,0,4,2,0,0,255] cc3_membered
s: 213
i: [1,216,0,4,2,0,0,255] cc5_membered
s: 216
i: [1,217,0,4,2,0,0,255] cc4_membered
s: 217
i: [1,222,0,4,2,0,0,255] cc1_membered
s: 222
d: [2,505,1,2,1,1,40,222,392]
a: [3,511,1,2,1,505]
s: 511
i: [1,225,0,4,2,0,0,255] cc4_xreal_0
s: 225
d: [2,506,1,2,1,1,40,222,402]
a: [3,516,1,2,1,506]
s: 516
i: [1,226,0,4,2,0,0,255] cc18_membered
s: 226
i: [1,261,0,4,1,0,0,3] existence_m1_subset_1
s: 261
d: [2,507,1,2,1,1,40,222,412]
a: [3,521,1,2,1,507]
s: 521
i: [1,227,0,5,2,0,0,255] t6_boole
s: 227
d: [2,529,1,2,1,1,40,225,377]
a: [3,536,1,2,1,529]
s: 536
i: [1,234,0,5,2,0,0,255] cc1_xxreal_0
s: 234
i: [1,256,0,5,2,0,0,255] t7_boole
s: 256
d: [2,544,1,2,1,1,40,226,355]
a: [3,551,1,2,1,544]
s: 551
a: [3,367,0,5,2,264]
s: 367
d: [2,545,1,2,1,1,40,226,382]
a: [3,556,1,2,1,545]
s: 556
i: [1,228,0,6,2,0,0,1] fc7_xxreal_1
s: 228
i: [1,235,0,6,2,0,0,255] cc7_membered
s: 235
d: [2,546,1,2,1,1,40,226,462]
a: [3,561,1,2,1,546]
s: 561
i: [1,241,0,6,3,0,0,255] cc7_xxreal_0
s: 241
d: [2,547,1,2,1,1,40,226,472]
a: [3,566,1,2,1,547]
s: 566
i: [1,242,0,6,3,0,0,255] cc7_xxreal_0
s: 242
i: [1,245,0,6,3,0,0,255] cc3_xxreal_0
s: 245
d: [2,608,1,4,2,1,40,241,442]
d: [2,610,1,2,1,2,21,608,447]
a: [3,614,1,2,1,610]
s: 614
i: [1,247,0,6,3,0,0,255] cc3_xxreal_0
s: 247
d: [2,609,1,4,2,1,40,241,477]
d: [2,616,1,2,1,2,21,609,482]
a: [3,620,1,2,1,616]
s: 620
i: [1,248,0,6,3,0,0,255] cc5_xxreal_0
s: 248
i: [1,250,0,6,3,0,0,255] cc5_xxreal_0
s: 250
d: [2,628,1,4,2,1,40,242,427]
d: [2,630,1,2,1,2,21,628,432]
a: [3,634,1,2,1,630]
s: 634
i: [1,251,0,6,2,0,0,255] t1_subset
s: 251
d: [2,629,1,4,2,1,40,242,487]
d: [2,636,1,2,1,2,21,629,492]
a: [3,640,1,2,1,636]
s: 640
i: [1,263,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 263
i: [1,218,0,7,3,0,0,255] cc14_membered
s: 218
d: [2,642,1,4,2,1,40,247,427]
d: [2,644,1,2,1,2,21,642,432]
a: [3,648,1,2,1,644]
s: 648
i: [1,220,0,7,3,0,0,2] cc13_membered
s: 220
d: [2,643,1,4,2,1,40,247,487]
d: [2,650,1,2,1,2,21,643,492]
a: [3,654,1,2,1,650]
s: 654
i: [1,223,0,7,3,0,0,255] cc17_membered
s: 223
i: [1,253,0,7,2,0,0,255] t3_subset
s: 253
d: [2,523,2,2,1,1,40,511,212]
a: [3,527,2,2,1,523]
s: 527
i: [1,254,0,7,2,0,0,255] t3_subset
s: 254
d: [2,538,2,2,1,1,40,516,212]
a: [3,542,2,2,1,538]
s: 542
i: [1,255,0,7,3,0,0,1] t8_boole
s: 255
i: [1,257,0,7,2,0,0,255] t2_tarski
s: 257
d: [2,568,2,2,1,1,40,521,212]
a: [3,572,2,2,1,568]
s: 572
i: [1,211,0,8,3,0,0,255] cc19_membered
s: 211
d: [2,582,2,2,1,1,40,551,222]
a: [3,587,2,2,1,582]
s: 587
i: [1,214,0,8,3,0,0,255] cc22_membered
s: 214
i: [1,215,0,8,3,0,0,255] cc23_membered
s: 215
d: [2,589,2,2,1,1,40,556,222]
a: [3,593,2,2,1,589]
s: 593
i: [1,219,0,8,3,0,0,255] cc21_membered
s: 219
d: [2,602,2,2,1,1,40,561,222]
a: [3,606,2,2,1,602]
s: 606
i: [1,221,0,8,3,0,0,2] cc20_membered
s: 221
i: [1,224,0,8,3,0,0,255] cc24_membered
s: 224
d: [2,622,2,2,1,1,40,566,222]
a: [3,626,2,2,1,622]
s: 626
i: [1,237,0,8,4,0,0,255] cc4_xxreal_0
s: 237
d: [2,723,1,6,3,1,40,237,335]
a: [3,738,1,2,1,723]
s: 738
r: 738
i: [1,239,0,8,4,0,0,255] cc6_xxreal_0
s: 239
i: [1,243,0,8,4,0,0,255] cc8_xxreal_0
s: 243
a: [3,737,1,2,1,723]
s: 737
i: [1,252,0,8,3,0,0,255] t2_subset
s: 252
a: [3,742,1,2,1,723]
s: 742
i: [1,265,0,8,3,0,0,1] fc3_xxreal_1
s: 265
i: [1,260,0,9,3,0,0,255] t5_subset
s: 260
d: [2,724,1,6,3,1,40,237,330]
a: [3,751,1,2,1,724]
s: 751
r: 751
i: [1,258,0,10,3,0,0,255] t2_tarski
s: 258
a: [3,750,1,2,1,724]
s: 750
i: [1,259,0,10,3,0,0,255] t4_subset
s: 259
i: [1,266,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 266
a: [3,755,1,2,1,724]
s: 755
i: [1,229,0,11,3,0,0,255] d3_xxreal_1
s: 229
d: [2,725,1,6,3,1,40,237,325]
a: [3,764,1,2,1,725]
s: 764
r: 764
i: [1,230,0,13,5,0,0,1] t1_xxreal_0
s: 230
i: [1,231,0,15,6,0,0,1] t2_xxreal_0
s: 231
a: [3,763,1,2,1,725]
s: 763
i: [1,267,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 267
a: [3,768,1,2,1,725]
s: 768
i: [1,268,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 268
i: [1,269,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 269
d: [2,726,1,6,3,1,40,237,320]
a: [3,777,1,2,1,726]
s: 777
r: 777
i: [1,270,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 270
a: [3,776,1,2,1,726]
s: 776
i: [1,297,0,18,6,0,0,255] fraenkel_a_2_2_xxreal_1
s: 297
i: [1,232,0,21,7,0,0,255] t61_xxreal_1
s: 232
a: [3,781,1,2,1,726]
s: 781
i: [1,233,0,21,7,0,0,255] t61_xxreal_1
s: 233
d: [2,841,1,7,3,1,40,266,310]
d: [2,844,1,5,2,2,21,841,330]
d: [2,845,1,3,1,3,21,844,335]
a: [3,849,1,3,1,845]
s: 849
d: [2,916,1,21,7,1,44,232,315]
d: [2,919,1,19,6,2,21,916,320]
d: [2,920,1,17,5,3,21,919,325]
a: [3,924,1,3,1,920]
s: 924
d: [2,822,1,8,3,1,44,265,315]
d: [2,823,1,6,2,2,21,822,320]
d: [2,824,1,4,1,3,21,823,325]
a: [3,828,1,4,1,824]
s: 828
d: [2,949,2,10,4,1,40,924,230]
d: [2,952,2,8,3,2,21,949,320]
d: [2,953,2,6,2,3,21,952,325]
a: [3,957,2,3,1,953]
s: 957
d: [2,728,1,6,3,1,40,237,387]
a: [3,803,1,2,1,728]
s: 803
r: 803
a: [3,802,1,2,1,728]
s: 802
a: [3,807,1,2,1,728]
s: 807
d: [2,868,1,16,5,0,40,267,223]
a: [3,875,1,2,1,868]
s: 875
d: [2,870,1,16,5,0,40,267,218]
a: [3,884,1,2,1,870]
s: 884
d: [2,576,1,3,1,0,40,234,261]
a: [3,580,1,3,1,576]
s: 580
d: [2,944,2,12,5,1,40,849,231]
d: [2,946,2,10,4,2,21,944,330]
d: [2,947,2,8,3,3,21,946,335]
s: 947
d: [2,596,1,4,1,0,40,235,261]
a: [3,600,1,4,1,596]
s: 600
d: [2,917,1,21,7,1,44,232,315]
d: [2,929,1,19,6,2,21,917,320]
d: [2,930,1,17,5,3,21,929,325]
s: 930
d: [2,992,2,10,4,3,40,930,262]
d: [2,994,2,7,3,4,21,992,310]
d: [2,995,2,5,2,5,21,994,330]
d: [2,996,2,3,1,6,21,995,335]
a: [3,1000,2,3,1,996]
s: 1000
d: [2,934,1,21,7,1,44,233,315]
d: [2,942,1,19,6,2,21,934,320]
d: [2,943,1,17,5,3,21,942,325]
s: 943
d: [2,1021,2,10,4,3,40,943,262]
d: [2,1023,2,7,3,4,21,1021,310]
d: [2,1024,2,5,2,5,21,1023,330]
d: [2,1025,2,3,1,6,21,1024,335]
a: [3,1029,2,3,1,1025]
s: 1029
d: [2,656,1,5,2,0,40,218,261]
s: 656
d: [2,657,1,5,2,0,40,220,261]
s: 657
d: [2,1002,3,5,2,4,40,1000,947]
d: [2,1005,3,3,1,5,21,1002,325]
a: [3,1009,3,3,1,1005]
s: 1009
d: [2,1044,4,0,0,6,52,325,330,320,1029,957,1009,231]
e: 1044
r: 924
r: 957
a: [3,923,1,3,1,920]
s: 923
a: [3,956,2,3,1,953]
s: 956
a: [3,927,1,14,4,920]
s: 927
r: 301
a: [3,300,0,3,1,203]
s: 300
r: 315
r: 325
r: 763
r: 768
r: 923
r: 927
r: 930
r: 956
r: 1000
r: 1009
a: [3,305,0,3,1,203]
s: 305
d: [2,1078,0,7,1,2,23,315,300]
a: [3,1091,0,7,1,1078]
s: 1091
d: [2,1082,1,3,1,2,23,923,300]
a: [3,1096,1,3,1,1082]
s: 1096
d: [2,933,1,21,7,1,44,233,315]
d: [2,936,1,19,6,2,21,933,320]
d: [2,937,1,17,5,3,21,936,325]
a: [3,940,1,14,4,937]
s: 940
d: [2,1032,3,10,4,1,40,1029,230]
d: [2,1035,3,8,3,2,21,1032,330]
d: [2,1036,3,6,2,3,21,1035,320]
d: [2,1069,2,7,3,1,40,927,262]
d: [2,1071,2,5,2,2,21,1069,335]
d: [2,1072,2,3,1,3,21,1071,330]
d: [2,1004,3,10,4,1,40,1000,230]
d: [2,1075,3,8,3,2,21,1004,325]
d: [2,1076,3,6,2,3,21,1075,335]
d: [2,1136,2,7,3,1,40,940,262]
d: [2,1138,2,5,2,2,21,1136,335]
d: [2,1139,2,3,1,3,21,1138,330]
f: 1139,1076,1072,1044,1036,1025,1005,996,953,937,920,845,197,198,199,200,201,202,203
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_1__t68_xxreal_1
% SZS output start Proof for xxreal_1__t68_xxreal_1
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ! [X3] : (v1_xxreal_0(X3) => (k3_xxreal_1(X0,X1) = k3_xxreal_1(X2,X3) => ((X1 = X3 & X0 = X2) | r1_xxreal_0(X1,X0))))))) [input t68_xxreal_1]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ! [X3] : (v1_xxreal_0(X3) => (k3_xxreal_1(X0,X1) = k3_xxreal_1(X2,X3) => ((X1 = X3 & X0 = X2) | r1_xxreal_0(X1,X0))))))) [negated conjecture 1]
35. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
64. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
67. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input t1_xxreal_0]
70. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input t2_xxreal_0]
74. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ! [X3] : (v1_xxreal_0(X3) => (r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(X2,X3)) => ((r1_xxreal_0(X1,X3) & r1_xxreal_0(X2,X0)) | r1_xxreal_0(X1,X0))))))) [input t61_xxreal_1]
86. ! [X0] : r1_tarski(X0,X0) [rectify 64]
96. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((X1 != X3 | X0 != X2) & ~r1_xxreal_0(X1,X0)) & k3_xxreal_1(X0,X1) = k3_xxreal_1(X2,X3)) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
97. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((X1 != X3 | X0 != X2) & ~r1_xxreal_0(X1,X0) & k3_xxreal_1(X0,X1) = k3_xxreal_1(X2,X3) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 96]
119. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 67]
120. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 119]
121. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 70]
122. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 121]
123. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((((r1_xxreal_0(X1,X3) & r1_xxreal_0(X2,X0)) | r1_xxreal_0(X1,X0)) | ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(X2,X3))) | ~v1_xxreal_0(X3)) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 74]
124. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_xxreal_0(X1,X3) & r1_xxreal_0(X2,X0)) | r1_xxreal_0(X1,X0) | ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(X2,X3)) | ~v1_xxreal_0(X3)) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 123]
153. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 35]
154. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 153]
157. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((X1 != X3 | X0 != X2) & ~r1_xxreal_0(X1,X0) & k3_xxreal_1(X0,X1) = k3_xxreal_1(X2,X3) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (? [X3] : ((X1 != X3 | sK0 != X2) & ~r1_xxreal_0(X1,sK0) & k3_xxreal_1(X2,X3) = k3_xxreal_1(sK0,X1) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
158. ? [X1] : (? [X2] : (? [X3] : ((X1 != X3 | sK0 != X2) & ~r1_xxreal_0(X1,sK0) & k3_xxreal_1(X2,X3) = k3_xxreal_1(sK0,X1) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (? [X3] : ((sK1 != X3 | sK0 != X2) & ~r1_xxreal_0(sK1,sK0) & k3_xxreal_1(X2,X3) = k3_xxreal_1(sK0,sK1) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
159. ? [X2] : (? [X3] : ((sK1 != X3 | sK0 != X2) & ~r1_xxreal_0(sK1,sK0) & k3_xxreal_1(X2,X3) = k3_xxreal_1(sK0,sK1) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) => (? [X3] : ((sK1 != X3 | sK0 != sK2) & ~r1_xxreal_0(sK1,sK0) & k3_xxreal_1(sK0,sK1) = k3_xxreal_1(sK2,X3) & v1_xxreal_0(X3)) & v1_xxreal_0(sK2)) [choice axiom]
160. ? [X3] : ((sK1 != X3 | sK0 != sK2) & ~r1_xxreal_0(sK1,sK0) & k3_xxreal_1(sK0,sK1) = k3_xxreal_1(sK2,X3) & v1_xxreal_0(X3)) => ((sK1 != sK3 | sK0 != sK2) & ~r1_xxreal_0(sK1,sK0) & k3_xxreal_1(sK0,sK1) = k3_xxreal_1(sK2,sK3) & v1_xxreal_0(sK3)) [choice axiom]
161. ((((sK1 != sK3 | sK0 != sK2) & ~r1_xxreal_0(sK1,sK0) & k3_xxreal_1(sK0,sK1) = k3_xxreal_1(sK2,sK3) & v1_xxreal_0(sK3)) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 97,160,159,158,157]
197. v1_xxreal_0(sK0) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
198. v1_xxreal_0(sK1) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
199. v1_xxreal_0(sK2) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
200. v1_xxreal_0(sK3) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
201. k3_xxreal_1(sK0,sK1) = k3_xxreal_1(sK2,sK3) [cnf transformation 161] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
202. ~r1_xxreal_0(sK1,sK0) [cnf transformation 161] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
203. sK1 != sK3 | sK0 != sK2 [cnf transformation 161] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
230. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 120] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
231. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 122] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
232. ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(X2,X3)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X2,X0) | ~v1_xxreal_0(X3) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 124] {a:0,w:21,wCS:8400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
233. ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(X2,X3)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X1,X3) | ~v1_xxreal_0(X3) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 124] {a:0,w:21,wCS:8400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
262. r1_tarski(X0,X0) [cnf transformation 86] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
266. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 154] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
299. 1 <=> sK0 = sK2 [avatar definition]
303. 2 <=> sK1 = sK3 [avatar definition]
306. ~1 | ~2 [avatar split clause 203,303,299]
308. 3 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
310. ~r1_xxreal_0(sK1,sK0) <- (~3) [avatar component clause 308] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
311. ~3 [avatar split clause 202,308]
313. 4 <=> k3_xxreal_1(sK0,sK1) = k3_xxreal_1(sK2,sK3) [avatar definition]
315. k3_xxreal_1(sK0,sK1) = k3_xxreal_1(sK2,sK3) <- (4) [avatar component clause 313] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
316. 4 [avatar split clause 201,313]
318. 5 <=> v1_xxreal_0(sK3) [avatar definition]
320. v1_xxreal_0(sK3) <- (5) [avatar component clause 318] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
321. 5 [avatar split clause 200,318]
323. 6 <=> v1_xxreal_0(sK2) [avatar definition]
325. v1_xxreal_0(sK2) <- (6) [avatar component clause 323] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
326. 6 [avatar split clause 199,323]
328. 7 <=> v1_xxreal_0(sK1) [avatar definition]
330. v1_xxreal_0(sK1) <- (7) [avatar component clause 328] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
331. 7 [avatar split clause 198,328]
333. 8 <=> v1_xxreal_0(sK0) [avatar definition]
335. v1_xxreal_0(sK0) <- (8) [avatar component clause 333] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
336. 8 [avatar split clause 197,333]
841. r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~3) [resolution 266,310] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
844. r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK0) <- (~3, 7) [subsumption resolution 841,330] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
845. r1_xxreal_0(sK0,sK1) <- (~3, 7, 8) [subsumption resolution 844,335] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
847. 92 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
849. r1_xxreal_0(sK0,sK1) <- (92) [avatar component clause 847] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
850. 92 | 3 | ~7 | ~8 [avatar split clause 845,333,328,308,847]
916. ~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) <- (4) [superposition 232,315] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
917. ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(sK2,X0) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4) [superposition 232,315] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
919. ~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK2) <- (4, 5) [subsumption resolution 916,320] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
920. ~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5, 6) [subsumption resolution 919,325] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
922. 99 <=> r1_xxreal_0(sK3,sK2) [avatar definition]
924. r1_xxreal_0(sK3,sK2) <- (99) [avatar component clause 922] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
926. 100 <=> ! [X1,X0] : (~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | r1_xxreal_0(X0,sK2)) [avatar definition]
927. ~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | r1_xxreal_0(X0,sK2) <- (100) [avatar component clause 926] {a:1,w:14,wCS:1400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
928. 99 | 100 | ~4 | ~5 | ~6 [avatar split clause 920,323,318,313,926,922]
929. ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(sK2,X0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5) [subsumption resolution 917,320] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
930. ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(sK2,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5, 6) [subsumption resolution 929,325] {a:1,w:17,wCS:1700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
933. ~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(sK3,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) <- (4) [superposition 233,315] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
934. ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X1,sK3) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4) [superposition 233,315] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
936. ~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(sK3,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK2) <- (4, 5) [subsumption resolution 933,320] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
937. ~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(sK3,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5, 6) [subsumption resolution 936,325] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
939. 101 <=> ! [X1,X0] : (~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | r1_xxreal_0(sK3,X1)) [avatar definition]
940. ~r1_tarski(k3_xxreal_1(sK0,sK1),k3_xxreal_1(X0,X1)) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | r1_xxreal_0(sK3,X1) <- (101) [avatar component clause 939] {a:1,w:14,wCS:1400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
941. 99 | 101 | ~4 | ~5 | ~6 [avatar split clause 937,323,318,313,939,922]
942. ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X1,sK3) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5) [subsumption resolution 934,320] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
943. ~r1_tarski(k3_xxreal_1(X0,X1),k3_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X1,sK3) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5, 6) [subsumption resolution 942,325] {a:1,w:17,wCS:1700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
944. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(X0) <- (92) [resolution 849,231] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
946. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(X0) <- (7, 92) [subsumption resolution 944,330] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
947. ~r1_xxreal_0(X0,sK0) | r1_xxreal_0(X0,sK1) | ~v1_xxreal_0(X0) <- (7, 8, 92) [subsumption resolution 946,335] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
949. sK2 = sK3 | ~r1_xxreal_0(sK2,sK3) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) <- (99) [resolution 924,230] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
952. sK2 = sK3 | ~r1_xxreal_0(sK2,sK3) | ~v1_xxreal_0(sK2) <- (5, 99) [subsumption resolution 949,320] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
953. sK2 = sK3 | ~r1_xxreal_0(sK2,sK3) <- (5, 6, 99) [subsumption resolution 952,325] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
955. 102 <=> r1_xxreal_0(sK2,sK3) [avatar definition]
957. ~r1_xxreal_0(sK2,sK3) <- (~102) [avatar component clause 955] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
959. 103 <=> sK2 = sK3 [avatar definition]
962. ~102 | 103 | ~5 | ~6 | ~99 [avatar split clause 953,922,323,318,959,955]
992. r1_xxreal_0(sK1,sK0) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4, 5, 6) [resolution 930,262] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
994. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~3, 4, 5, 6) [subsumption resolution 992,310] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
995. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) <- (~3, 4, 5, 6, 7) [subsumption resolution 994,330] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
996. r1_xxreal_0(sK2,sK0) <- (~3, 4, 5, 6, 7, 8) [subsumption resolution 995,335] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
998. 108 <=> r1_xxreal_0(sK2,sK0) [avatar definition]
1000. r1_xxreal_0(sK2,sK0) <- (108) [avatar component clause 998] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1001. 108 | 3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 996,333,328,323,318,313,308,998]
1002. r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK2) <- (7, 8, 92, 108) [resolution 1000,947] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
1004. sK0 = sK2 | ~r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (108) [resolution 1000,230] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1005. r1_xxreal_0(sK2,sK1) <- (6, 7, 8, 92, 108) [subsumption resolution 1002,325] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
1007. 109 <=> r1_xxreal_0(sK2,sK1) [avatar definition]
1009. r1_xxreal_0(sK2,sK1) <- (109) [avatar component clause 1007] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1010. 109 | ~6 | ~7 | ~8 | ~92 | ~108 [avatar split clause 1005,998,847,333,328,323,1007]
1017. 110 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
1021. r1_xxreal_0(sK1,sK0) | r1_xxreal_0(sK1,sK3) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4, 5, 6) [resolution 943,262] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1023. r1_xxreal_0(sK1,sK3) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~3, 4, 5, 6) [subsumption resolution 1021,310] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1024. r1_xxreal_0(sK1,sK3) | ~v1_xxreal_0(sK0) <- (~3, 4, 5, 6, 7) [subsumption resolution 1023,330] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1025. r1_xxreal_0(sK1,sK3) <- (~3, 4, 5, 6, 7, 8) [subsumption resolution 1024,335] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1027. 111 <=> r1_xxreal_0(sK1,sK3) [avatar definition]
1029. r1_xxreal_0(sK1,sK3) <- (111) [avatar component clause 1027] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1030. 111 | 3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 1025,333,328,323,318,313,308,1027]
1032. sK1 = sK3 | ~r1_xxreal_0(sK3,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK3) <- (111) [resolution 1029,230] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1035. sK1 = sK3 | ~r1_xxreal_0(sK3,sK1) | ~v1_xxreal_0(sK3) <- (7, 111) [subsumption resolution 1032,330] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1036. sK1 = sK3 | ~r1_xxreal_0(sK3,sK1) <- (5, 7, 111) [subsumption resolution 1035,320] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1038. 112 <=> r1_xxreal_0(sK3,sK1) [avatar definition]
1041. ~112 | 2 | ~5 | ~7 | ~111 [avatar split clause 1036,1027,328,318,303,1038]
1044. $false <- (5, 6, 7, ~102, 109, 111) [unit resulting resolution 325,330,320,1029,957,1009,231] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:44,thDist:-44}
1047. ~5 | ~6 | ~7 | 102 | ~109 | ~111 [avatar contradiction clause 1044]
1069. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | r1_xxreal_0(sK0,sK2) <- (100) [resolution 927,262] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1071. ~v1_xxreal_0(sK1) | r1_xxreal_0(sK0,sK2) <- (8, 100) [subsumption resolution 1069,335] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1072. r1_xxreal_0(sK0,sK2) <- (7, 8, 100) [subsumption resolution 1071,330] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1073. 110 | ~7 | ~8 | ~100 [avatar split clause 1072,926,333,328,1017]
1074. sK2 != sK3 | r1_xxreal_0(sK3,sK1) | ~r1_xxreal_0(sK2,sK1) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
1075. sK0 = sK2 | ~r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK0) <- (6, 108) [subsumption resolution 1004,325] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1076. sK0 = sK2 | ~r1_xxreal_0(sK0,sK2) <- (6, 8, 108) [subsumption resolution 1075,335] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1077. ~110 | 1 | ~6 | ~8 | ~108 [avatar split clause 1076,998,333,323,299,1017]
1108. sK1 != sK3 | sK2 != sK3 | r1_xxreal_0(sK1,sK0) | ~r1_xxreal_0(sK2,sK0) [theory tautology sat conflict] {a:0,w:12,wCS:4800,sine:255,thAx:0,allAx:1,thDist:-1}
1136. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | r1_xxreal_0(sK3,sK1) <- (101) [resolution 940,262] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1138. ~v1_xxreal_0(sK1) | r1_xxreal_0(sK3,sK1) <- (8, 101) [subsumption resolution 1136,335] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1139. r1_xxreal_0(sK3,sK1) <- (7, 8, 101) [subsumption resolution 1138,330] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1140. 112 | ~7 | ~8 | ~101 [avatar split clause 1139,939,333,328,1038]
1141. $false [avatar sat refutation 306,311,316,321,326,331,336,850,928,941,962,1001,1010,1030,1041,1047,1073,1074,1077,1108,1140] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:768,thDist:-768}
% SZS output end Proof for xxreal_1__t68_xxreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 76
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 100
% Generated clauses: 474
% Active clauses: 150
% Passive clauses: 229
% Final active clauses: 130
% Final passive clauses: 72
% Main loop iterations started: 155
% 
% >>> Simplifying Inferences
% Duplicate literals: 33
% Fw subsumption resolutions: 66
% Fw demodulations: 1
% Bw demodulations: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 1
% Forward subsumptions: 42
% 
% >>> Generating Inferences
% Binary resolution: 136
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 14
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 19
% Split components: 45
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 115
% SAT solver unit clauses: 38
% SAT solver binary clauses: 36
% 
% Memory used [KB]: 5628
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% backward subsumption index maintenance: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

