i: [1,55,0,2,1,1,0,0] t23_pboole
a: [3,99,0,2,1,55]
s: 99
i: [1,53,0,2,1,1,0,0] t23_pboole
a: [3,109,0,2,1,53]
s: 109
i: [1,51,0,2,1,1,0,0] t23_pboole
a: [3,119,0,2,1,51]
s: 119
i: [1,49,0,2,1,1,0,0] t23_pboole
a: [3,129,0,2,1,49]
s: 129
i: [1,56,0,3,1,1,0,0] t23_pboole
a: [3,94,0,3,1,56]
s: 94
i: [1,54,0,3,1,1,0,0] t23_pboole
a: [3,104,0,3,1,54]
s: 104
i: [1,52,0,3,1,1,0,0] t23_pboole
a: [3,114,0,3,1,52]
s: 114
i: [1,50,0,3,1,1,0,0] t23_pboole
a: [3,124,0,3,1,50]
s: 124
i: [1,57,0,4,1,1,0,0] t23_pboole
a: [3,89,0,4,1,57]
s: 89
i: [1,58,0,7,1,1,0,0] t23_pboole
a: [3,84,0,7,1,58]
s: 84
i: [1,78,0,2,1,0,0,1] rc1_funct_1
a: [3,134,0,2,1,78]
s: 134
i: [1,77,0,2,1,0,0,1] rc1_funct_1
a: [3,139,0,2,1,77]
s: 139
i: [1,62,0,3,1,0,0,1] rc1_pboole
s: 62
i: [1,64,0,3,1,0,0,1] rc1_pboole
s: 64
i: [1,63,0,4,1,0,0,1] rc1_pboole
s: 63
i: [1,65,0,4,1,0,0,1] rc1_pboole
s: 65
i: [1,79,0,24,9,0,0,1] redefinition_r6_pboole
d: [2,80,0,14,5,0,18,79]
s: 80
i: [1,66,0,24,9,0,0,1] reflexivity_r6_pboole
s: 66
i: [1,67,0,24,9,0,0,1] reflexivity_r2_pboole
s: 67
i: [1,70,0,25,9,0,0,1] dt_k3_pboole
s: 70
i: [1,72,0,25,9,0,0,1] dt_k3_pboole
s: 72
i: [1,68,0,26,9,0,0,2] idempotence_k3_pboole
s: 68
i: [1,71,0,26,9,0,0,1] dt_k3_pboole
s: 71
i: [1,73,0,26,9,0,0,1] dt_k3_pboole
s: 73
i: [1,59,0,27,9,0,0,1] t15_pboole
s: 59
i: [1,75,0,27,10,0,0,1] redefinition_r6_pboole
s: 75
i: [1,74,0,28,10,0,0,1] symmetry_r6_pboole
s: 74
i: [1,69,0,29,9,0,0,2] commutativity_k3_pboole
s: 69
d: [2,160,1,47,17,0,44,59,69]
d: [2,169,1,27,9,0,18,160]
s: 169
i: [1,60,0,31,11,0,0,2] l22_pboole
s: 60
i: [1,61,0,45,15,0,0,1] t17_pboole
s: 61
d: [2,181,1,27,10,1,40,60,89]
d: [2,188,1,24,9,2,21,181,114]
d: [2,189,1,22,8,3,21,188,119]
d: [2,190,1,19,7,4,21,189,124]
d: [2,191,1,17,6,5,21,190,129]
d: [2,192,1,14,5,6,21,191,94]
d: [2,193,1,12,4,7,21,192,99]
d: [2,194,1,9,3,8,21,193,104]
d: [2,195,1,7,2,9,21,194,109]
a: [3,199,1,4,1,195]
s: 199
d: [2,155,1,47,17,0,44,59,69]
d: [2,174,1,27,9,0,18,155]
s: 174
d: [2,183,1,65,18,0,40,60,59]
d: [2,186,1,55,14,0,18,183]
d: [2,205,1,49,13,0,21,186,73]
d: [2,206,1,44,12,0,21,205,72]
d: [2,207,1,38,11,0,21,206,71]
d: [2,208,1,33,10,0,21,207,70]
s: 208
d: [2,213,1,83,24,0,40,61,60]
d: [2,220,1,73,20,0,18,213]
d: [2,221,1,67,19,0,21,220,73]
d: [2,222,1,62,18,0,21,221,72]
d: [2,223,1,56,17,0,21,222,71]
d: [2,224,1,51,16,0,21,223,70]
s: 224
d: [2,225,2,64,23,0,40,208,61]
d: [2,231,2,34,11,0,18,225]
d: [2,232,2,30,10,0,21,231,67]
s: 232
d: [2,258,3,28,10,1,44,84,232]
d: [2,282,3,25,9,2,21,258,114]
d: [2,283,3,23,8,3,21,282,119]
d: [2,284,3,20,7,4,21,283,124]
d: [2,285,3,18,6,5,21,284,129]
d: [2,286,3,15,5,6,21,285,94]
d: [2,287,3,13,4,7,21,286,99]
d: [2,288,3,10,3,8,21,287,104]
d: [2,289,3,8,2,9,21,288,109]
d: [2,290,3,4,1,10,21,289,89]
a: [3,294,3,4,1,290]
s: 294
d: [2,296,4,0,0,5,52,129,119,124,114,294,80]
e: 296
d: [2,297,4,10,4,1,40,294,80]
d: [2,299,4,7,3,2,21,297,114]
d: [2,300,4,5,2,3,21,299,119]
d: [2,301,4,2,1,4,21,300,124]
d: [2,302,4,0,0,5,21,301,129]
e: 302
f: 302,290,49,50,51,52,53,54,55,56,57,58
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pboole__t23_pboole
% SZS output start Proof for pboole__t23_pboole
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r2_pboole(X0,X1,X2) => r6_pboole(X0,k3_pboole(X0,X1,X2),X1)))) [input t23_pboole]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r2_pboole(X0,X1,X2) => r6_pboole(X0,k3_pboole(X0,X1,X2),X1)))) [negated conjecture 1]
4. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2)))) [input dt_k3_pboole]
6. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => X1 = X2))) [input l22_pboole]
9. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (r6_pboole(X0,X1,X2) <=> X1 = X2)) [input redefinition_r6_pboole]
10. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => r2_pboole(X0,X1,X1)) [input reflexivity_r2_pboole]
13. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k3_pboole(X0,X1,X2),X1))) [input t15_pboole]
14. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X1,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,k3_pboole(X0,X2,X3)))))) [input t17_pboole]
15. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r2_pboole(X0,X1,X2) => r6_pboole(X0,k3_pboole(X0,X1,X2),X1)))) [flattening 2]
16. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k3_pboole(X0,X1,X2),X1))) [flattening 13]
17. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => X1 = X2))) [flattening 6]
18. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X1,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,k3_pboole(X0,X2,X3)))))) [flattening 14]
19. ? [X0,X1] : (? [X2] : ((~r6_pboole(X0,k3_pboole(X0,X1,X2),X1) & r2_pboole(X0,X1,X2)) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 15]
20. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k3_pboole(X0,X1,X2),X1) & r2_pboole(X0,X1,X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 19]
21. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 16]
22. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 21]
23. ! [X0,X1] : (! [X2] : ((X1 = X2 | (~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 17]
24. ! [X0,X1] : (! [X2] : (X1 = X2 | ~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 23]
25. ! [X0,X1] : (! [X2] : (! [X3] : ((r2_pboole(X0,X1,k3_pboole(X0,X2,X3)) | (~r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 18]
26. ! [X0,X1] : (! [X2] : (! [X3] : (r2_pboole(X0,X1,k3_pboole(X0,X2,X3)) | ~r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 25]
29. ! [X0,X1,X2] : (r2_pboole(X0,X1,X1) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 10]
30. ! [X0,X1,X2] : (r2_pboole(X0,X1,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 29]
35. ! [X0,X1,X2] : ((v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 4]
36. ! [X0,X1,X2] : ((v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 35]
39. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 9]
40. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 39]
41. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k3_pboole(X0,X1,X2),X1) & r2_pboole(X0,X1,X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (~r6_pboole(sK0,k3_pboole(sK0,sK1,X2),sK1) & r2_pboole(sK0,sK1,X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
42. ? [X2] : (~r6_pboole(sK0,k3_pboole(sK0,sK1,X2),sK1) & r2_pboole(sK0,sK1,X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (~r6_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
43. (~r6_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 20,42,41]
46. ! [X0,X1,X2] : (((r6_pboole(X0,X1,X2) | X1 != X2) & (X1 = X2 | ~r6_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 40]
49. v1_relat_1(sK1) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. v4_relat_1(sK1,sK0) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. v1_funct_1(sK1) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. v1_partfun1(sK1,sK0) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
53. v1_relat_1(sK2) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. v4_relat_1(sK2,sK0) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. v1_funct_1(sK2) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. v1_partfun1(sK2,sK0) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
57. r2_pboole(sK0,sK1,sK2) [cnf transformation 43] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. ~r6_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) [cnf transformation 43] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
59. r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 22] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. ~r2_pboole(X0,X2,X1) | X1 = X2 | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 24] {a:0,w:31,wCS:12400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
61. r2_pboole(X0,X1,k3_pboole(X0,X2,X3)) | ~r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 26] {a:0,w:45,wCS:18000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. r2_pboole(X0,X1,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 30] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. v1_relat_1(k3_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 36] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. v4_relat_1(k3_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 36] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. v1_funct_1(k3_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 36] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. v1_partfun1(k3_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 36] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
76. r6_pboole(X0,X1,X2) | X1 != X2 | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 46] {a:0,w:27,wCS:10800,sine:1,thAx:0,allAx:1,thDist:-1}
79. r6_pboole(X0,X2,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) [equality resolution 76] {a:0,w:24,wCS:9600,sine:1,thAx:0,allAx:1,thDist:-1}
80. r6_pboole(X0,X2,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) [duplicate literal removal 79] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
82. 1 <=> r6_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) [avatar definition]
84. ~r6_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) <- (~1) [avatar component clause 82] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
85. ~1 [avatar split clause 58,82]
87. 2 <=> r2_pboole(sK0,sK1,sK2) [avatar definition]
89. r2_pboole(sK0,sK1,sK2) <- (2) [avatar component clause 87] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
90. 2 [avatar split clause 57,87]
92. 3 <=> v1_partfun1(sK2,sK0) [avatar definition]
94. v1_partfun1(sK2,sK0) <- (3) [avatar component clause 92] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. 3 [avatar split clause 56,92]
97. 4 <=> v1_funct_1(sK2) [avatar definition]
99. v1_funct_1(sK2) <- (4) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. 4 [avatar split clause 55,97]
102. 5 <=> v4_relat_1(sK2,sK0) [avatar definition]
104. v4_relat_1(sK2,sK0) <- (5) [avatar component clause 102] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 5 [avatar split clause 54,102]
107. 6 <=> v1_relat_1(sK2) [avatar definition]
109. v1_relat_1(sK2) <- (6) [avatar component clause 107] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. 6 [avatar split clause 53,107]
112. 7 <=> v1_partfun1(sK1,sK0) [avatar definition]
114. v1_partfun1(sK1,sK0) <- (7) [avatar component clause 112] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
115. 7 [avatar split clause 52,112]
117. 8 <=> v1_funct_1(sK1) [avatar definition]
119. v1_funct_1(sK1) <- (8) [avatar component clause 117] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
120. 8 [avatar split clause 51,117]
122. 9 <=> v4_relat_1(sK1,sK0) [avatar definition]
124. v4_relat_1(sK1,sK0) <- (9) [avatar component clause 122] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
125. 9 [avatar split clause 50,122]
127. 10 <=> v1_relat_1(sK1) [avatar definition]
129. v1_relat_1(sK1) <- (10) [avatar component clause 127] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
130. 10 [avatar split clause 49,127]
183. k3_pboole(X4,X3,X5) = X3 | ~r2_pboole(X4,X3,k3_pboole(X4,X3,X5)) | ~v1_partfun1(k3_pboole(X4,X3,X5),X4) | ~v1_funct_1(k3_pboole(X4,X3,X5)) | ~v4_relat_1(k3_pboole(X4,X3,X5),X4) | ~v1_relat_1(k3_pboole(X4,X3,X5)) | ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) [resolution 60,59] {a:1,w:65,wCS:26000,sine:1,thAx:0,allAx:2,thDist:-2}
186. k3_pboole(X4,X3,X5) = X3 | ~r2_pboole(X4,X3,k3_pboole(X4,X3,X5)) | ~v1_partfun1(k3_pboole(X4,X3,X5),X4) | ~v1_funct_1(k3_pboole(X4,X3,X5)) | ~v4_relat_1(k3_pboole(X4,X3,X5),X4) | ~v1_relat_1(k3_pboole(X4,X3,X5)) | ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) [duplicate literal removal 183] {a:1,w:55,wCS:22000,sine:1,thAx:0,allAx:2,thDist:-2}
205. k3_pboole(X4,X3,X5) = X3 | ~r2_pboole(X4,X3,k3_pboole(X4,X3,X5)) | ~v1_funct_1(k3_pboole(X4,X3,X5)) | ~v4_relat_1(k3_pboole(X4,X3,X5),X4) | ~v1_relat_1(k3_pboole(X4,X3,X5)) | ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) [subsumption resolution 186,73] {a:1,w:49,wCS:19600,sine:1,thAx:0,allAx:2,thDist:-2}
206. k3_pboole(X4,X3,X5) = X3 | ~r2_pboole(X4,X3,k3_pboole(X4,X3,X5)) | ~v4_relat_1(k3_pboole(X4,X3,X5),X4) | ~v1_relat_1(k3_pboole(X4,X3,X5)) | ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) [subsumption resolution 205,72] {a:1,w:44,wCS:17600,sine:1,thAx:0,allAx:2,thDist:-2}
207. k3_pboole(X4,X3,X5) = X3 | ~r2_pboole(X4,X3,k3_pboole(X4,X3,X5)) | ~v1_relat_1(k3_pboole(X4,X3,X5)) | ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) [subsumption resolution 206,71] {a:1,w:38,wCS:15200,sine:1,thAx:0,allAx:2,thDist:-2}
208. ~r2_pboole(X4,X3,k3_pboole(X4,X3,X5)) | k3_pboole(X4,X3,X5) = X3 | ~v1_partfun1(X3,X4) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X4) | ~v1_relat_1(X3) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) [subsumption resolution 207,70] {a:1,w:33,wCS:13200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
225. k3_pboole(X0,X1,X2) = X1 | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~r2_pboole(X0,X1,X2) | ~r2_pboole(X0,X1,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [resolution 208,61] {a:2,w:64,wCS:25600,sine:1,thAx:0,allAx:3,thDist:-3}
231. k3_pboole(X0,X1,X2) = X1 | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~r2_pboole(X0,X1,X2) | ~r2_pboole(X0,X1,X1) [duplicate literal removal 225] {a:2,w:34,wCS:13600,sine:1,thAx:0,allAx:3,thDist:-3}
232. k3_pboole(X0,X1,X2) = X1 | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~r2_pboole(X0,X1,X2) [subsumption resolution 231,67] {a:2,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
258. ~r6_pboole(sK0,sK1,sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~r2_pboole(sK0,sK1,sK2) <- (~1) [superposition 84,232] {a:3,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
282. ~r6_pboole(sK0,sK1,sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~r2_pboole(sK0,sK1,sK2) <- (~1, 7) [subsumption resolution 258,114] {a:3,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
283. ~r6_pboole(sK0,sK1,sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~r2_pboole(sK0,sK1,sK2) <- (~1, 7, 8) [subsumption resolution 282,119] {a:3,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
284. ~r6_pboole(sK0,sK1,sK1) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~r2_pboole(sK0,sK1,sK2) <- (~1, 7, 8, 9) [subsumption resolution 283,124] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
285. ~r6_pboole(sK0,sK1,sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~r2_pboole(sK0,sK1,sK2) <- (~1, 7, 8, 9, 10) [subsumption resolution 284,129] {a:3,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
286. ~r6_pboole(sK0,sK1,sK1) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~r2_pboole(sK0,sK1,sK2) <- (~1, 3, 7, 8, 9, 10) [subsumption resolution 285,94] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
287. ~r6_pboole(sK0,sK1,sK1) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~r2_pboole(sK0,sK1,sK2) <- (~1, 3, 4, 7, 8, 9, 10) [subsumption resolution 286,99] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
288. ~r6_pboole(sK0,sK1,sK1) | ~v1_relat_1(sK2) | ~r2_pboole(sK0,sK1,sK2) <- (~1, 3, 4, 5, 7, 8, 9, 10) [subsumption resolution 287,104] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
289. ~r6_pboole(sK0,sK1,sK1) | ~r2_pboole(sK0,sK1,sK2) <- (~1, 3, 4, 5, 6, 7, 8, 9, 10) [subsumption resolution 288,109] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
290. ~r6_pboole(sK0,sK1,sK1) <- (~1, 2, 3, 4, 5, 6, 7, 8, 9, 10) [subsumption resolution 289,89] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
292. 15 <=> r6_pboole(sK0,sK1,sK1) [avatar definition]
294. ~r6_pboole(sK0,sK1,sK1) <- (~15) [avatar component clause 292] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
295. ~15 | 1 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 [avatar split clause 290,127,122,117,112,107,102,97,92,87,82,292]
297. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~15) [resolution 294,80] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
299. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, ~15) [subsumption resolution 297,114] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
300. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (7, 8, ~15) [subsumption resolution 299,119] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
301. ~v1_relat_1(sK1) <- (7, 8, 9, ~15) [subsumption resolution 300,124] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
302. $false <- (7, 8, 9, 10, ~15) [subsumption resolution 301,129] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
303. ~7 | ~8 | ~9 | ~10 | 15 [avatar contradiction clause 302]
304. $false [avatar sat refutation 85,90,95,100,105,110,115,120,125,130,295,303] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:104,thDist:-104}
% SZS output end Proof for pboole__t23_pboole
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 30
% Generated clauses: 190
% Active clauses: 35
% Passive clauses: 47
% Final active clauses: 35
% Final passive clauses: 10
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Duplicate literals: 386
% Fw subsumption resolutions: 44
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Equational tautologies: 3
% Forward subsumptions: 22
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 1
% Forward superposition: 16
% Backward superposition: 23
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 12
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.003 s

