i: [1,161,0,2,1,1,0,0] l3_polyeq_2
a: [3,221,0,2,1,161]
s: 221
i: [1,162,0,2,1,0,0,2] spc2_boole
a: [3,226,0,2,1,162]
s: 226
i: [1,163,0,2,1,0,0,1] spc4_boole
a: [3,231,0,2,1,163]
s: 231
i: [1,164,0,2,1,0,0,5] fc3_membered
a: [3,236,0,2,1,164]
s: 236
i: [1,165,0,2,1,0,0,255] fc6_membered
a: [3,241,0,2,1,165]
s: 241
i: [1,210,0,2,1,0,0,7] rc2_xcmplx_0
a: [3,303,0,2,1,210]
s: 303
i: [1,212,0,2,1,0,0,4] rc2_membered
a: [3,308,0,2,1,212]
s: 308
i: [1,211,0,2,1,0,0,4] rc2_membered
a: [3,313,0,2,1,211]
s: 313
i: [1,214,0,2,1,0,0,4] rc1_membered
a: [3,318,0,2,1,214]
s: 318
i: [1,213,0,2,1,0,0,4] rc1_membered
a: [3,323,0,2,1,213]
s: 323
i: [1,216,0,2,1,0,0,255] rc3_membered
a: [3,328,0,2,1,216]
s: 328
i: [1,215,0,2,1,0,0,255] rc3_membered
a: [3,333,0,2,1,215]
s: 333
i: [1,217,0,2,1,0,0,255] rc4_xreal_0
a: [3,338,0,2,1,217]
s: 338
i: [1,203,0,3,1,0,0,255] reflexivity_r1_tarski
s: 203
i: [1,166,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,246,0,3,1,166]
s: 246
r: 241
d: [2,340,0,2,1,2,23,241,246]
a: [3,344,0,2,1,340]
s: 344
i: [1,167,0,3,1,0,0,2] l2_polyeq_2
a: [3,251,0,3,1,167]
s: 251
i: [1,168,0,3,1,0,0,255] t1_numerals
a: [3,256,0,3,1,168]
s: 256
i: [1,172,0,3,1,0,0,255] spc4_numerals
a: [3,266,0,3,1,172]
s: 266
i: [1,171,0,3,1,0,0,255] spc4_numerals
a: [3,271,0,3,1,171]
s: 271
i: [1,175,0,3,1,0,0,255] spc2_numerals
a: [3,281,0,3,1,175]
s: 281
i: [1,174,0,3,1,0,0,255] spc2_numerals
a: [3,286,0,3,1,174]
s: 286
d: [2,346,0,3,1,2,22,256,246]
a: [3,350,0,3,1,346]
s: 350
i: [1,180,0,4,2,0,0,3] cc2_membered
s: 180
i: [1,181,0,4,2,0,0,4] cc3_membered
s: 181
i: [1,185,0,4,2,0,0,255] cc4_membered
s: 185
d: [2,352,1,2,1,1,40,185,236]
a: [3,356,1,2,1,352]
s: 356
i: [1,186,0,4,2,0,0,255] cc5_membered
s: 186
d: [2,358,1,2,1,1,40,186,236]
a: [3,362,1,2,1,358]
s: 362
i: [1,188,0,4,2,0,0,3] cc1_membered
s: 188
i: [1,190,0,4,2,0,0,4] cc18_membered
s: 190
d: [2,364,1,2,1,1,40,188,344]
a: [3,371,1,2,1,364]
s: 371
i: [1,202,0,4,1,0,0,3] existence_m1_subset_1
s: 202
d: [2,365,1,2,1,1,40,188,308]
a: [3,376,1,2,1,365]
s: 376
i: [1,169,0,4,1,0,0,255] dt_k5_numbers
a: [3,261,0,4,1,169]
s: 261
i: [1,170,0,4,1,0,0,255] spc4_numerals
a: [3,276,0,4,1,170]
s: 276
d: [2,366,1,2,1,1,40,188,318]
a: [3,381,1,2,1,366]
s: 381
i: [1,173,0,4,1,0,0,255] spc2_numerals
a: [3,291,0,4,1,173]
s: 291
d: [2,367,1,2,1,1,40,188,328]
a: [3,386,1,2,1,367]
s: 386
i: [1,191,0,5,2,0,0,255] t6_boole
s: 191
i: [1,192,0,5,2,0,0,255] cc25_membered
s: 192
d: [2,388,1,2,1,1,40,190,338]
a: [3,392,1,2,1,388]
s: 392
i: [1,199,0,5,2,0,0,255] t7_boole
s: 199
d: [2,420,1,2,1,1,40,192,271]
a: [3,427,1,2,1,420]
s: 427
i: [1,209,0,7,3,0,0,2] reflexivity_r1_int_1
a: [3,298,0,5,2,209]
s: 298
i: [1,193,0,6,2,0,0,255] cc11_membered
s: 193
d: [2,421,1,2,1,1,40,192,350]
a: [3,432,1,2,1,421]
s: 432
i: [1,194,0,6,2,0,0,5] cc8_membered
s: 194
d: [2,422,1,2,1,1,40,192,286]
a: [3,437,1,2,1,422]
s: 437
i: [1,195,0,6,2,0,0,255] t1_subset
s: 195
i: [1,204,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 204
d: [2,469,1,2,1,1,40,194,261]
a: [3,474,1,2,1,469]
s: 474
i: [1,176,0,7,3,0,0,1] d1_abian
s: 176
d: [2,394,2,2,1,1,40,371,180]
a: [3,398,2,2,1,394]
s: 398
i: [1,177,0,7,3,0,0,1] d1_abian
s: 177
i: [1,182,0,7,3,0,0,2] cc16_membered
s: 182
d: [2,500,1,4,2,1,40,177,251]
d: [2,505,1,2,1,2,21,500,221]
d: [2,521,1,4,2,1,40,182,271]
d: [2,532,1,2,1,2,21,521,371]
f: 532,505,364,340,171,167,166,165,161
% Refutation found. Thanks to Tanya!
% SZS status Theorem for polyeq_2__l3_polyeq_2
% SZS output start Proof for polyeq_2__l3_polyeq_2
1. v1_abian(4) [input l3_polyeq_2]
2. ~v1_abian(4) [negated conjecture 1]
9. ! [X0] : (v5_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_int_1(X1))) [input cc16_membered]
25. ! [X0] : (v6_membered(X0) => v5_membered(X0)) [input cc1_membered]
34. v6_membered(k4_ordinal1) [input fc6_membered]
49. k5_numbers = k4_ordinal1 [input redefinition_k5_numbers]
67. m1_subset_1(4,k1_numbers) & m1_subset_1(4,k5_numbers) & m2_subset_1(4,k1_numbers,k5_numbers) & v2_xxreal_0(4) [input spc4_numerals]
70. r1_int_1(2,4) [input l2_polyeq_2]
71. ! [X0] : (v1_int_1(X0) => (v1_abian(X0) <=> r1_int_1(2,X0))) [input d1_abian]
72. ~v1_abian(4) [flattening 2]
87. m1_subset_1(4,k1_numbers) & m1_subset_1(4,k5_numbers) & m2_subset_1(4,k1_numbers,k5_numbers) [pure predicate removal 67]
109. ! [X0] : ((v1_abian(X0) <=> r1_int_1(2,X0)) | ~v1_int_1(X0)) [ennf transformation 71]
114. ! [X0] : (! [X1] : (v1_int_1(X1) | ~m1_subset_1(X1,X0)) | ~v5_membered(X0)) [ennf transformation 9]
120. ! [X0] : (v5_membered(X0) | ~v6_membered(X0)) [ennf transformation 25]
145. ! [X0] : (((v1_abian(X0) | ~r1_int_1(2,X0)) & (r1_int_1(2,X0) | ~v1_abian(X0))) | ~v1_int_1(X0)) [nnf transformation 109]
161. ~v1_abian(4) [cnf transformation 72] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
165. v6_membered(k4_ordinal1) [cnf transformation 34] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
166. k5_numbers = k4_ordinal1 [cnf transformation 49] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
167. r1_int_1(2,4) [cnf transformation 70] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
171. m1_subset_1(4,k5_numbers) [cnf transformation 87] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
177. ~r1_int_1(2,X0) | v1_abian(X0) | ~v1_int_1(X0) [cnf transformation 145] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
182. ~m1_subset_1(X1,X0) | v1_int_1(X1) | ~v5_membered(X0) [cnf transformation 114] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
188. ~v6_membered(X0) | v5_membered(X0) [cnf transformation 120] {a:0,w:4,wCS:1600,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
219. 1 <=> v1_abian(4) [avatar definition]
221. ~v1_abian(4) <- (~1) [avatar component clause 219] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
222. ~1 [avatar split clause 161,219]
239. 5 <=> v6_membered(k4_ordinal1) [avatar definition]
241. v6_membered(k4_ordinal1) <- (5) [avatar component clause 239] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
242. 5 [avatar split clause 165,239]
244. 6 <=> k5_numbers = k4_ordinal1 [avatar definition]
246. k5_numbers = k4_ordinal1 <- (6) [avatar component clause 244] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
247. 6 [avatar split clause 166,244]
249. 7 <=> r1_int_1(2,4) [avatar definition]
251. r1_int_1(2,4) <- (7) [avatar component clause 249] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
252. 7 [avatar split clause 167,249]
269. 11 <=> m1_subset_1(4,k5_numbers) [avatar definition]
271. m1_subset_1(4,k5_numbers) <- (11) [avatar component clause 269] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
272. 11 [avatar split clause 171,269]
340. v6_membered(k5_numbers) <- (5, 6) [backward demodulation 241,246] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
342. 26 <=> v6_membered(k5_numbers) [avatar definition]
344. v6_membered(k5_numbers) <- (26) [avatar component clause 342] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
345. 26 | ~5 | ~6 [avatar split clause 340,244,239,342]
364. v5_membered(k5_numbers) <- (26) [resolution 188,344] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
369. 30 <=> v5_membered(k5_numbers) [avatar definition]
371. v5_membered(k5_numbers) <- (30) [avatar component clause 369] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
372. 30 | ~26 [avatar split clause 364,342,369]
500. v1_abian(4) | ~v1_int_1(4) <- (7) [resolution 177,251] {a:1,w:4,wCS:1600,sine:1,thAx:0,allAx:2,thDist:-2}
505. ~v1_int_1(4) <- (~1, 7) [subsumption resolution 500,221] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
507. 52 <=> v1_int_1(4) [avatar definition]
510. ~52 | 1 | ~7 [avatar split clause 505,249,219,507]
521. v1_int_1(4) | ~v5_membered(k5_numbers) <- (11) [resolution 182,271] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:2,thDist:-2}
532. v1_int_1(4) <- (11, 30) [subsumption resolution 521,371] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
533. 52 | ~11 | ~30 [avatar split clause 532,369,269,507]
552. $false [avatar sat refutation 222,242,247,252,272,345,372,510,533] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:219,thDist:-219}
% SZS output end Proof for polyeq_2__l3_polyeq_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 70
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 7
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 57
% Generated clauses: 155
% Active clauses: 55
% Passive clauses: 88
% Final active clauses: 54
% Final passive clauses: 32
% Main loop iterations started: 56
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 4
% Fw demodulations: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 35
% Forward superposition: 1
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 60
% SAT solver unit clauses: 26
% SAT solver binary clauses: 22
% 
% Memory used [KB]: 9594
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% superposition: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

