i: [1,49,0,9,3,1,0,0] t3_xboole_0
a: [3,79,0,3,1,49]
s: 79
i: [1,47,0,9,3,1,0,0] t3_xboole_0
a: [3,84,0,3,1,47]
s: 84
i: [1,45,0,9,3,1,0,0] t3_xboole_0
a: [3,90,0,3,1,45]
s: 90
i: [1,50,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,96,0,2,1,50]
s: 96
i: [1,51,0,2,1,0,0,2] fc1_xboole_0
a: [3,101,0,2,1,51]
s: 101
i: [1,68,0,2,1,0,0,2] rc2_xboole_0
a: [3,111,0,2,1,68]
s: 111
i: [1,69,0,2,1,0,0,2] rc1_xboole_0
a: [3,116,0,2,1,69]
s: 116
i: [1,52,0,3,1,0,0,255] d2_xboole_0
a: [3,106,0,3,1,52]
s: 106
r: 96
i: [1,53,0,5,2,0,0,2] l13_xboole_0
s: 53
i: [1,54,0,5,2,0,0,1] d1_xboole_0
s: 54
i: [1,64,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 64
d: [2,121,1,2,1,1,40,54,84]
a: [3,126,1,2,1,121]
s: 126
i: [1,55,0,6,2,0,0,1] d1_xboole_0
s: 55
d: [2,122,1,2,1,1,40,54,90]
a: [3,131,1,2,1,122]
s: 131
i: [1,66,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 66
r: 79
a: [3,78,0,3,1,49]
s: 78
a: [3,75,0,6,2,49]
s: 75
d: [2,140,1,0,0,3,52,90,84,75]
e: 140
d: [2,141,1,3,1,2,40,75,84]
d: [2,144,1,0,0,3,21,141,90]
e: 144
r: 90
a: [3,89,0,3,1,45]
s: 89
d: [2,142,1,6,2,1,40,75,55]
d: [2,146,1,4,1,2,21,142,126]
a: [3,151,1,4,1,146]
s: 151
i: [1,67,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 67
i: [1,65,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 65
d: [2,153,1,3,1,1,40,67,84]
a: [3,158,1,3,1,153]
s: 158
i: [1,56,0,8,2,0,0,1] d7_xboole_0
s: 56
d: [2,154,1,6,2,0,40,67,55]
s: 154
i: [1,57,0,8,2,0,0,1] d7_xboole_0
s: 57
i: [1,71,0,8,2,0,0,1] d4_xboole_0
s: 71
d: [2,160,1,6,2,0,44,56,64]
s: 160
i: [1,72,0,8,2,0,0,1] d4_xboole_0
s: 72
d: [2,163,1,6,2,0,44,64,56]
s: 163
i: [1,70,0,11,3,0,0,1] d4_xboole_0
s: 70
i: [1,61,0,17,3,0,0,1] d4_xboole_0
s: 61
d: [2,193,1,17,3,1,40,61,75]
s: 193
i: [1,62,0,17,3,0,0,1] d4_xboole_0
s: 62
d: [2,199,1,17,3,1,40,62,75]
s: 199
i: [1,63,0,23,4,0,0,1] d4_xboole_0
s: 63
d: [2,166,1,6,2,0,44,57,64]
s: 166
d: [2,207,2,3,1,0,47,166]
a: [3,211,2,3,1,207]
s: 211
d: [2,119,1,7,3,0,44,53,53]
s: 119
d: [2,215,3,5,2,1,44,211,53]
s: 215
d: [2,161,1,8,2,0,44,56,65]
s: 161
d: [2,162,1,8,2,0,44,56,65]
s: 162
d: [2,164,1,8,2,0,44,65,56]
s: 164
d: [2,165,1,8,2,0,44,65,56]
s: 165
d: [2,168,1,8,2,0,44,57,65]
s: 168
d: [2,169,1,8,2,0,44,57,65]
s: 169
d: [2,173,1,8,2,0,44,71,65]
s: 173
d: [2,174,1,8,2,0,44,71,65]
s: 174
d: [2,172,1,9,3,0,44,71,56]
s: 172
d: [2,206,2,8,3,0,44,166,53]
s: 206
d: [2,177,1,9,3,0,44,72,56]
s: 177
d: [2,170,1,10,2,0,40,71,55]
s: 170
d: [2,250,2,10,2,1,40,170,75]
s: 250
d: [2,175,1,10,2,0,40,72,55]
s: 175
r: 250
r: 199
r: 193
r: 75
r: 78
r: 89
s: 79
s: 90
d: [2,134,1,3,1,1,40,66,79]
a: [3,138,1,3,1,134]
s: 138
d: [2,282,1,3,1,1,40,90,67]
a: [3,287,1,3,1,282]
s: 287
d: [2,183,1,10,3,0,40,70,54]
s: 183
d: [2,245,2,6,2,0,40,170,54]
s: 245
d: [2,182,1,11,3,0,40,70,67]
s: 182
d: [2,194,1,11,2,0,42,61]
s: 194
d: [2,263,2,6,2,0,40,175,54]
s: 263
d: [2,200,1,11,2,0,42,62]
s: 200
d: [2,302,3,6,2,0,44,245,65]
s: 302
d: [2,185,1,12,4,0,44,70,56]
s: 185
r: 79
r: 84
r: 90
s: 79
s: 84
s: 90
d: [2,336,2,9,3,1,40,185,84]
a: [3,351,2,3,1,336]
s: 351
d: [2,367,3,0,0,2,52,101,351,54]
e: 367
d: [2,372,3,2,1,1,40,351,54]
d: [2,375,3,0,0,2,21,372,101]
e: 375
r: 351
d: [2,190,1,13,3,0,40,61,54]
s: 190
d: [2,196,1,13,3,0,40,62,54]
s: 196
a: [3,350,2,3,1,336]
s: 350
d: [2,396,3,0,0,4,52,79,90,84,350,185]
e: 396
d: [2,397,3,0,0,4,52,138,84,90,350,185]
e: 397
r: 79
r: 90
r: 138
s: 78
s: 89
s: 75
a: [3,137,1,3,1,134]
s: 137
d: [2,261,3,8,2,1,40,175,250]
d: [2,274,3,4,1,1,18,261]
a: [3,278,3,4,1,274]
s: 278
d: [2,405,1,17,3,1,40,75,62]
s: 405
a: [3,347,2,6,2,336]
s: 347
d: [2,406,1,17,3,1,40,75,61]
s: 406
d: [2,189,1,17,3,0,40,61,67]
s: 189
d: [2,418,3,3,1,2,40,347,84]
a: [3,423,3,3,1,418]
s: 423
d: [2,195,1,17,3,0,40,62,67]
s: 195
d: [2,337,2,9,3,1,40,185,90]
a: [3,355,2,6,2,337]
s: 355
d: [2,191,1,23,3,0,40,61,72]
s: 191
d: [2,192,1,23,3,0,40,61,71]
s: 192
d: [2,403,2,10,2,1,40,75,175]
s: 403
d: [2,197,1,23,3,0,40,62,72]
s: 197
d: [2,404,2,10,2,1,40,75,170]
s: 404
d: [2,198,1,23,3,0,40,62,71]
s: 198
d: [2,203,1,39,5,0,40,63,70]
s: 203
d: [2,464,3,10,2,1,44,403,65]
s: 464
d: [2,407,2,11,2,1,40,75,194]
s: 407
d: [2,465,3,10,2,1,44,403,65]
s: 465
d: [2,409,2,11,2,1,40,75,200]
s: 409
d: [2,408,2,13,3,1,40,75,190]
s: 408
d: [2,527,3,14,4,1,40,408,196]
d: [2,540,3,7,2,1,18,527]
d: [2,542,3,7,2,1,22,540,65]
s: 542
d: [2,410,2,13,3,1,40,75,196]
s: 410
d: [2,549,4,8,3,1,44,57,542]
d: [2,567,4,5,2,2,21,549,78]
d: [2,568,4,2,1,2,21,567,53]
s: 568
r: 111
r: 126
r: 131
r: 183
r: 542
r: 410
r: 408
r: 263
r: 245
r: 215
r: 206
r: 196
r: 190
r: 119
r: 54
r: 53
d: [2,590,5,0,0,3,40,568,278]
e: 590
d: [2,589,5,0,0,3,40,568,101]
e: 589
d: [2,588,5,0,0,3,40,568,116]
e: 588
d: [2,587,5,0,0,3,52,116,568]
e: 587
i: [1,46,0,6,2,1,0,0] t3_xboole_0
i: [1,44,0,6,2,1,0,0] t3_xboole_0
f: 587,396,375,141,69,51,44,45,46,47,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_0__t3_xboole_0
% SZS output start Proof for xboole_0__t3_xboole_0
1. ! [X0] : ! [X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X2] : ~(r2_hidden(X2,X1) & r2_hidden(X2,X0)) & ~r1_xboole_0(X0,X1))) [input t3_xboole_0]
2. ~! [X0] : ! [X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X2] : ~(r2_hidden(X2,X1) & r2_hidden(X2,X0)) & ~r1_xboole_0(X0,X1))) [negated conjecture 1]
4. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
5. ! [X0] : (v1_xboole_0(X0) <=> ! [X1] : ~r2_hidden(X1,X0)) [input d1_xboole_0]
7. ! [X0] : ! [X1] : ! [X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [input d4_xboole_0]
8. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) <=> k3_xboole_0(X0,X1) = k1_xboole_0) [input d7_xboole_0]
12. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
14. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input l13_xboole_0]
15. ? [X0] : v1_xboole_0(X0) [input rc1_xboole_0]
18. ~! [X0] : ! [X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X3] : ~(r2_hidden(X3,X1) & r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) [rectify 2]
19. ~! [X0,X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X3] : ~(r2_hidden(X3,X1) & r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) [flattening 18]
20. ! [X0,X1] : (r1_xboole_0(X0,X1) <=> k3_xboole_0(X0,X1) = k1_xboole_0) [flattening 8]
21. ! [X0,X1,X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [flattening 7]
23. ? [X0,X1] : ((r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) | (! [X3] : (~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) [ennf transformation 19]
24. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 14]
27. ? [X0,X1] : ((r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) | (! [X3] : (~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) => ((r1_xboole_0(sK0,sK1) & ? [X2] : (r2_hidden(X2,sK1) & r2_hidden(X2,sK0))) | (! [X3] : (~r2_hidden(X3,sK1) | ~r2_hidden(X3,sK0)) & ~r1_xboole_0(sK0,sK1))) [choice axiom]
28. ? [X2] : (r2_hidden(X2,sK1) & r2_hidden(X2,sK0)) => (r2_hidden(sK2,sK1) & r2_hidden(sK2,sK0)) [choice axiom]
29. (r1_xboole_0(sK0,sK1) & (r2_hidden(sK2,sK1) & r2_hidden(sK2,sK0))) | (! [X3] : (~r2_hidden(X3,sK1) | ~r2_hidden(X3,sK0)) & ~r1_xboole_0(sK0,sK1)) [skolemisation 23,28,27]
30. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X1] : ~r2_hidden(X1,X0) | ~v1_xboole_0(X0))) [nnf transformation 5]
31. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [rectify 30]
32. ! [X0] : (? [X1] : r2_hidden(X1,X0) => r2_hidden(sK3(X0),X0)) [choice axiom]
33. ! [X0] : ((v1_xboole_0(X0) | r2_hidden(sK3(X0),X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [skolemisation 31,32]
34. ! [X0,X1] : ((r1_xboole_0(X0,X1) | k3_xboole_0(X0,X1) != k1_xboole_0) & (k3_xboole_0(X0,X1) = k1_xboole_0 | ~r1_xboole_0(X0,X1))) [nnf transformation 20]
35. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) | ~r2_hidden(X3,X0))) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [nnf transformation 21]
36. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | ~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [flattening 35]
37. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [rectify 36]
38. ! [X2,X1,X0] : (? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2))) => ((~r2_hidden(sK4(X0,X1,X2),X1) | ~r2_hidden(sK4(X0,X1,X2),X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & ((r2_hidden(sK4(X0,X1,X2),X1) & r2_hidden(sK4(X0,X1,X2),X0)) | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
39. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ((~r2_hidden(sK4(X0,X1,X2),X1) | ~r2_hidden(sK4(X0,X1,X2),X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & ((r2_hidden(sK4(X0,X1,X2),X1) & r2_hidden(sK4(X0,X1,X2),X0)) | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [skolemisation 37,38]
42. ? [X0] : v1_xboole_0(X0) => v1_xboole_0(sK6) [choice axiom]
43. v1_xboole_0(sK6) [skolemisation 15,42]
44. r2_hidden(sK2,sK0) | ~r1_xboole_0(sK0,sK1) [cnf transformation 29] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
45. r2_hidden(sK2,sK0) | ~r2_hidden(X3,sK1) | ~r2_hidden(X3,sK0) [cnf transformation 29] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
46. r2_hidden(sK2,sK1) | ~r1_xboole_0(sK0,sK1) [cnf transformation 29] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
47. r2_hidden(sK2,sK1) | ~r2_hidden(X3,sK1) | ~r2_hidden(X3,sK0) [cnf transformation 29] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. r1_xboole_0(sK0,sK1) | ~r2_hidden(X3,sK1) | ~r2_hidden(X3,sK0) [cnf transformation 29] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. v1_xboole_0(k1_xboole_0) [cnf transformation 12] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
53. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 24] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
54. ~r2_hidden(X2,X0) | ~v1_xboole_0(X0) [cnf transformation 33] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
56. k3_xboole_0(X0,X1) = k1_xboole_0 | ~r1_xboole_0(X0,X1) [cnf transformation 34] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. k3_xboole_0(X0,X1) != k1_xboole_0 | r1_xboole_0(X0,X1) [cnf transformation 34] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) | k3_xboole_0(X0,X1) != X2 [cnf transformation 39] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
61. r2_hidden(sK4(X0,X1,X2),X2) | r2_hidden(sK4(X0,X1,X2),X0) | k3_xboole_0(X0,X1) = X2 [cnf transformation 39] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
62. r2_hidden(sK4(X0,X1,X2),X2) | r2_hidden(sK4(X0,X1,X2),X1) | k3_xboole_0(X0,X1) = X2 [cnf transformation 39] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
65. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
69. v1_xboole_0(sK6) [cnf transformation 43] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
70. r2_hidden(X4,k3_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) [equality resolution 60] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
74. 1 <=> ! [X3] : (~r2_hidden(X3,sK1) | ~r2_hidden(X3,sK0)) [avatar definition]
75. ~r2_hidden(X3,sK1) | ~r2_hidden(X3,sK0) <- (1) [avatar component clause 74] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
77. 2 <=> r1_xboole_0(sK0,sK1) [avatar definition]
78. ~r1_xboole_0(sK0,sK1) <- (~2) [avatar component clause 77] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
79. r1_xboole_0(sK0,sK1) <- (2) [avatar component clause 77] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
80. 1 | 2 [avatar split clause 49,77,74]
82. 3 <=> r2_hidden(sK2,sK1) [avatar definition]
84. r2_hidden(sK2,sK1) <- (3) [avatar component clause 82] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
85. 1 | 3 [avatar split clause 47,82,74]
86. ~2 | 3 [avatar split clause 46,82,77]
88. 4 <=> r2_hidden(sK2,sK0) [avatar definition]
90. r2_hidden(sK2,sK0) <- (4) [avatar component clause 88] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. 1 | 4 [avatar split clause 45,88,74]
92. ~2 | 4 [avatar split clause 44,88,77]
99. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
101. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 99] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
102. 6 [avatar split clause 51,99]
114. 9 <=> v1_xboole_0(sK6) [avatar definition]
116. v1_xboole_0(sK6) <- (9) [avatar component clause 114] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
117. 9 [avatar split clause 69,114]
141. ~r2_hidden(sK2,sK0) <- (1, 3) [resolution 75,84] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
147. ~4 | ~1 | ~3 [avatar split clause 141,82,74,88]
185. ~r2_hidden(X4,X3) | r2_hidden(X4,k1_xboole_0) | ~r2_hidden(X4,X2) | ~r1_xboole_0(X2,X3) [superposition 70,56] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
190. r2_hidden(sK4(X3,X4,X5),X3) | k3_xboole_0(X3,X4) = X5 | ~v1_xboole_0(X5) [resolution 61,54] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
196. r2_hidden(sK4(X3,X4,X5),X4) | k3_xboole_0(X3,X4) = X5 | ~v1_xboole_0(X5) [resolution 62,54] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
349. 19 <=> r2_hidden(sK2,k1_xboole_0) [avatar definition]
350. ~r2_hidden(sK2,k1_xboole_0) <- (~19) [avatar component clause 349] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
351. r2_hidden(sK2,k1_xboole_0) <- (19) [avatar component clause 349] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
372. ~v1_xboole_0(k1_xboole_0) <- (19) [resolution 351,54] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
375. $false <- (6, 19) [subsumption resolution 372,101] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
376. ~6 | ~19 [avatar contradiction clause 375]
396. $false <- (2, 3, 4, ~19) [unit resulting resolution 79,90,84,350,185] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
399. ~2 | ~3 | ~4 | 19 [avatar contradiction clause 396]
408. ~r2_hidden(sK4(sK1,X7,X8),sK0) | k3_xboole_0(sK1,X7) = X8 | ~v1_xboole_0(X8) <- (1) [resolution 75,190] {a:2,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
527. k3_xboole_0(sK1,sK0) = X0 | ~v1_xboole_0(X0) | k3_xboole_0(sK1,sK0) = X0 | ~v1_xboole_0(X0) <- (1) [resolution 408,196] {a:3,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
540. k3_xboole_0(sK1,sK0) = X0 | ~v1_xboole_0(X0) <- (1) [duplicate literal removal 527] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
542. k3_xboole_0(sK0,sK1) = X0 | ~v1_xboole_0(X0) <- (1) [forward demodulation 540,65] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
549. k1_xboole_0 != X3 | r1_xboole_0(sK0,sK1) | ~v1_xboole_0(X3) <- (1) [superposition 57,542] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
567. k1_xboole_0 != X3 | ~v1_xboole_0(X3) <- (1, ~2) [subsumption resolution 549,78] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
568. ~v1_xboole_0(X3) <- (1, ~2) [subsumption resolution 567,53] {a:4,w:2,wCS:200,nSel:1,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
587. $false <- (1, ~2, 9) [unit resulting resolution 116,568] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
594. ~1 | 2 | ~9 [avatar contradiction clause 587]
595. $false [avatar sat refutation 80,85,86,91,92,102,117,147,376,399,594] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:296,thDist:-296}
% SZS output end Proof for xboole_0__t3_xboole_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 471
% Active clauses: 89
% Passive clauses: 221
% Final active clauses: 57
% Final passive clauses: 115
% Main loop iterations started: 99
% 
% >>> Simplifying Inferences
% Duplicate literals: 44
% Trivial inequalities: 2
% Fw subsumption resolutions: 12
% Fw demodulations: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 24
% Equational tautologies: 1
% Forward subsumptions: 136
% Backward subsumptions: 16
% 
% >>> Generating Inferences
% Binary resolution: 193
% Unit resulting resolution: 5
% Factoring: 6
% Forward superposition: 107
% Backward superposition: 31
% Self superposition: 2
% Equality resolution: 2
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 18
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 38
% SAT solver unit clauses: 6
% SAT solver binary clauses: 15
% 
% Memory used [KB]: 9722
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% term sharing: 0.001 s
% forward subsumption: 0.002 s
% resolution: 0.001 s ( own 0.0000 s ) 
% superposition: 0.001 s
% other: 0.007 s ( own 0.002 s ) 

