i: [1,115,0,2,1,1,0,0] t31_ordinal3
a: [3,183,0,2,1,115]
s: 183
i: [1,114,0,2,1,1,0,0] t31_ordinal3
a: [3,188,0,2,1,114]
s: 188
i: [1,118,0,3,1,1,0,0] t31_ordinal3
a: [3,168,0,3,1,118]
s: 168
i: [1,117,0,3,1,1,0,0] t31_ordinal3
a: [3,173,0,3,1,117]
s: 173
i: [1,116,0,5,1,1,0,0] t31_ordinal3
a: [3,178,0,5,1,116]
s: 178
i: [1,119,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,193,0,2,1,119]
s: 193
i: [1,120,0,2,1,0,0,1] fc1_xboole_0
a: [3,198,0,2,1,120]
s: 198
i: [1,155,0,2,1,0,0,2] rc2_xboole_0
a: [3,215,0,2,1,155]
s: 215
i: [1,156,0,2,1,0,0,255] rc4_ordinal1
a: [3,220,0,2,1,156]
s: 220
i: [1,157,0,2,1,0,0,2] rc1_xboole_0
a: [3,225,0,2,1,157]
s: 225
i: [1,158,0,2,1,0,0,1] rc1_ordinal1
a: [3,230,0,2,1,158]
s: 230
i: [1,162,0,2,1,0,0,255] rc2_ordinal1
a: [3,235,0,2,1,162]
s: 235
i: [1,161,0,2,1,0,0,255] rc2_ordinal1
a: [3,240,0,2,1,161]
s: 240
i: [1,160,0,2,1,0,0,255] rc2_ordinal1
a: [3,245,0,2,1,160]
s: 245
i: [1,159,0,2,1,0,0,255] rc2_ordinal1
a: [3,250,0,2,1,159]
s: 250
i: [1,164,0,2,1,0,0,255] rc5_ordinal1
a: [3,255,0,2,1,164]
s: 255
i: [1,163,0,2,1,0,0,255] rc5_ordinal1
a: [3,260,0,2,1,163]
s: 260
i: [1,122,0,3,1,0,0,255] t2_xboole_1
s: 122
i: [1,146,0,3,1,0,0,255] irreflexivity_r2_xboole_0
s: 146
i: [1,147,0,3,1,0,0,255] reflexivity_r1_tarski
s: 147
i: [1,121,0,3,1,0,0,255] d2_xboole_0
a: [3,203,0,3,1,121]
s: 203
r: 193
i: [1,125,0,4,2,0,0,255] cc6_ordinal1
s: 125
i: [1,126,0,4,2,0,0,255] cc9_ordinal1
s: 126
d: [2,263,1,2,1,1,40,125,220]
a: [3,268,1,2,1,263]
s: 268
i: [1,127,0,4,2,0,0,255] cc7_ordinal1
s: 127
i: [1,128,0,4,2,0,0,1] cc3_ordinal1
s: 128
d: [2,264,1,2,1,1,40,125,255]
a: [3,273,1,2,1,264]
s: 273
i: [1,130,0,4,2,0,0,255] cc1_ordinal1
s: 130
d: [2,311,1,2,1,1,40,130,188]
a: [3,320,1,2,1,311]
s: 320
i: [1,131,0,4,2,0,0,255] cc1_ordinal1
s: 131
i: [1,145,0,4,1,0,0,3] existence_m1_subset_1
s: 145
d: [2,312,1,2,1,1,40,130,183]
a: [3,325,1,2,1,312]
s: 325
i: [1,129,0,5,2,0,0,1] t6_boole
s: 129
d: [2,342,1,2,1,1,40,131,188]
a: [3,351,1,2,1,342]
s: 351
i: [1,142,0,5,2,0,0,2] t7_boole
s: 142
i: [1,150,0,7,3,0,0,255] reflexivity_r1_ordinal1
a: [3,210,0,5,2,150]
s: 210
d: [2,343,1,2,1,1,40,131,183]
a: [3,356,1,2,1,343]
s: 356
i: [1,135,0,6,3,0,0,255] cc2_ordinal1
s: 135
d: [2,275,1,2,1,1,40,126,225]
a: [3,280,1,2,1,275]
s: 280
i: [1,137,0,6,2,0,0,2] t1_subset
s: 137
i: [1,148,0,6,2,0,0,255] antisymmetry_r2_xboole_0
s: 148
d: [2,276,1,2,1,1,40,126,198]
a: [3,285,1,2,1,276]
s: 285
i: [1,149,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 149
d: [2,287,1,2,1,1,40,127,225]
a: [3,292,1,2,1,287]
s: 292
i: [1,133,0,7,3,0,0,3] cc5_ordinal1
s: 133
i: [1,134,0,7,3,0,0,3] cc2_ordinal2
s: 134
d: [2,288,1,2,1,1,40,127,198]
a: [3,297,1,2,1,288]
s: 297
i: [1,139,0,7,2,0,0,255] t3_subset
s: 139
d: [2,299,1,2,1,1,40,128,225]
a: [3,304,1,2,1,299]
s: 304
i: [1,140,0,7,2,0,0,255] t3_subset
s: 140
i: [1,141,0,7,3,0,0,1] t8_boole
s: 141
d: [2,300,1,2,1,1,40,128,198]
a: [3,309,1,2,1,300]
s: 309
i: [1,123,0,8,3,0,0,255] cc10_ordinal1
s: 123
d: [2,314,1,2,1,1,40,130,230]
a: [3,335,1,2,1,314]
s: 335
i: [1,136,0,8,3,0,0,2] t2_subset
s: 136
i: [1,151,0,8,3,0,0,1] dt_k11_ordinal2
s: 151
d: [2,345,1,2,1,1,40,131,230]
a: [3,366,1,2,1,345]
s: 366
i: [1,138,0,9,3,0,0,255] d8_xboole_0
s: 138
d: [2,313,2,2,1,1,40,130,268]
a: [3,330,2,2,1,313]
s: 330
i: [1,144,0,9,3,0,0,255] t5_subset
s: 144
i: [1,124,0,10,4,0,0,255] t11_ordinal1
s: 124
d: [2,316,2,2,1,1,40,130,273]
a: [3,340,2,2,1,316]
s: 340
i: [1,143,0,10,3,0,0,255] t4_subset
s: 143
d: [2,344,2,2,1,1,40,131,268]
a: [3,361,2,2,1,344]
s: 361
i: [1,152,0,10,4,0,0,255] connectedness_r1_ordinal1
s: 152
i: [1,153,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 153
d: [2,347,2,2,1,1,40,131,273]
a: [3,371,2,2,1,347]
s: 371
i: [1,154,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 154
d: [2,387,2,2,1,1,40,304,131]
a: [3,392,2,2,1,387]
s: 392
i: [1,132,0,19,6,0,0,1] t40_ordinal2
s: 132
d: [2,444,1,17,6,1,44,132,178]
d: [2,446,1,14,5,2,21,444,168]
d: [2,447,1,12,4,3,21,446,188]
d: [2,448,1,10,3,4,21,447,183]
s: 448
d: [2,456,2,8,3,4,44,448,178]
d: [2,457,2,6,2,4,21,456,188]
d: [2,458,2,3,1,4,21,457,149]
a: [3,462,2,3,1,458]
s: 462
d: [2,445,1,17,6,1,44,132,178]
d: [2,449,1,14,5,2,21,445,168]
d: [2,450,1,12,4,3,21,449,188]
d: [2,451,1,10,3,4,21,450,183]
s: 451
d: [2,468,2,12,4,4,44,451,129]
d: [2,470,2,7,3,4,21,468,142]
a: [3,473,2,2,1,470]
s: 473
r: 260
r: 250
r: 215
r: 144
r: 142
r: 141
r: 129
r: 128
r: 127
r: 126
d: [2,480,3,0,0,2,40,473,198]
e: 480
d: [2,479,3,0,0,2,40,473,225]
e: 479
d: [2,478,3,0,0,2,52,225,473]
e: 478
r: 473
s: 215
s: 250
d: [2,466,2,7,3,4,40,451,142]
d: [2,469,2,5,2,5,21,466,198]
s: 469
s: 260
a: [3,476,2,5,2,470]
s: 476
r: 451
s: 126
s: 127
d: [2,388,2,2,1,1,40,304,130]
a: [3,397,2,2,1,388]
s: 397
s: 128
d: [2,401,2,2,1,1,40,309,131]
a: [3,407,2,2,1,401]
s: 407
s: 129
s: 142
d: [2,402,2,2,1,1,40,309,130]
a: [3,412,2,2,1,402]
s: 412
s: 141
d: [2,454,2,9,3,4,40,448,142]
s: 454
s: 144
d: [2,374,1,5,2,0,44,122,129]
s: 374
d: [2,501,2,9,3,4,40,142,448]
s: 501
d: [2,382,1,5,2,0,40,133,145]
s: 382
d: [2,452,2,10,3,4,40,448,149]
s: 452
d: [2,386,1,5,1,0,40,139,145]
s: 386
d: [2,436,1,7,3,0,40,154,122]
d: [2,439,1,5,2,1,21,436,309]
s: 439
d: [2,453,2,10,3,4,40,448,137]
s: 453
d: [2,493,1,5,2,0,44,122,129]
s: 493
d: [2,464,2,10,3,4,40,451,149]
s: 464
d: [2,414,1,6,2,0,40,123,145]
s: 414
d: [2,416,1,6,2,0,40,136,145]
s: 416
d: [2,525,3,5,2,1,40,416,476]
a: [3,529,3,3,1,525]
s: 529
d: [2,535,4,0,0,2,52,188,529,382]
e: 535
d: [2,536,4,0,0,2,52,188,145,529,133]
e: 536
d: [2,537,4,2,1,1,40,529,382]
d: [2,540,4,0,0,2,21,537,188]
e: 540
r: 529
d: [2,422,1,6,2,0,40,138,122]
s: 422
a: [3,533,3,2,1,525]
s: 533
d: [2,547,4,0,0,2,52,173,533,129]
e: 547
d: [2,548,4,0,0,3,52,198,173,533,141]
e: 548
d: [2,549,4,0,0,3,52,198,173,533,141]
e: 549
f: 549,540,525,469,120,114,115,116,117,118
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t31_ordinal3
% SZS output start Proof for ordinal3__t31_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ~(k1_xboole_0 != X1 & k1_xboole_0 != X0 & k11_ordinal2(X0,X1) = k1_xboole_0))) [input t31_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ~(k1_xboole_0 != X1 & k1_xboole_0 != X0 & k11_ordinal2(X0,X1) = k1_xboole_0))) [negated conjecture 1]
11. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (m1_subset_1(X1,X0) => v3_ordinal1(X1))) [input cc5_ordinal1]
23. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
24. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
37. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
40. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r2_hidden(X1,X2) => (r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0))))) [input t40_ordinal2]
44. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
45. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [input t8_boole]
46. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 37]
50. ! [X0,X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [flattening 45]
51. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 44]
58. ? [X0] : (? [X1] : ((k1_xboole_0 != X1 & k1_xboole_0 != X0 & k11_ordinal2(X0,X1) = k1_xboole_0) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
59. ? [X0] : (? [X1] : (k1_xboole_0 != X1 & k1_xboole_0 != X0 & k11_ordinal2(X0,X1) = k1_xboole_0 & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 58]
69. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0) | ~r2_hidden(X1,X2)) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 40]
70. ! [X0] : (! [X1] : (! [X2] : (r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0 | ~r2_hidden(X1,X2) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 69]
71. ! [X0] : (! [X1] : (v3_ordinal1(X1) | ~m1_subset_1(X1,X0)) | ~v3_ordinal1(X0)) [ennf transformation 11]
75. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 46]
76. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 75]
80. ! [X0,X1] : (~v1_xboole_0(X1) | X0 = X1 | ~v1_xboole_0(X0)) [ennf transformation 50]
81. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 51]
95. ? [X0] : (? [X1] : (k1_xboole_0 != X1 & k1_xboole_0 != X0 & k11_ordinal2(X0,X1) = k1_xboole_0 & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (k1_xboole_0 != X1 & k1_xboole_0 != sK0 & k1_xboole_0 = k11_ordinal2(sK0,X1) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
96. ? [X1] : (k1_xboole_0 != X1 & k1_xboole_0 != sK0 & k1_xboole_0 = k11_ordinal2(sK0,X1) & v3_ordinal1(X1)) => (k1_xboole_0 != sK1 & k1_xboole_0 != sK0 & k1_xboole_0 = k11_ordinal2(sK0,sK1) & v3_ordinal1(sK1)) [choice axiom]
97. (k1_xboole_0 != sK1 & k1_xboole_0 != sK0 & k1_xboole_0 = k11_ordinal2(sK0,sK1) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 59,96,95]
99. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK2(X0),X0)) [choice axiom]
100. ! [X0] : m1_subset_1(sK2(X0),X0) [skolemisation 23,99]
114. v3_ordinal1(sK0) [cnf transformation 97] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
115. v3_ordinal1(sK1) [cnf transformation 97] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
116. k1_xboole_0 = k11_ordinal2(sK0,sK1) [cnf transformation 97] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
117. k1_xboole_0 != sK0 [cnf transformation 97] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
118. k1_xboole_0 != sK1 [cnf transformation 97] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
120. v1_xboole_0(k1_xboole_0) [cnf transformation 24] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
132. r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0 | ~r2_hidden(X1,X2) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 70] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
133. ~m1_subset_1(X1,X0) | v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 71] {a:0,w:7,wCS:2800,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
136. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 76] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
141. X0 = X1 | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 80] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
142. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 81] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
145. m1_subset_1(sK2(X0),X0) [cnf transformation 100] {a:0,w:4,wCS:1600,nSel:1,sine:3,thAx:0,allAx:2,thDist:-2}
166. 1 <=> k1_xboole_0 = sK1 [avatar definition]
168. k1_xboole_0 != sK1 <- (~1) [avatar component clause 166] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
169. ~1 [avatar split clause 118,166]
171. 2 <=> k1_xboole_0 = sK0 [avatar definition]
173. k1_xboole_0 != sK0 <- (~2) [avatar component clause 171] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
174. ~2 [avatar split clause 117,171]
176. 3 <=> k1_xboole_0 = k11_ordinal2(sK0,sK1) [avatar definition]
178. k1_xboole_0 = k11_ordinal2(sK0,sK1) <- (3) [avatar component clause 176] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
179. 3 [avatar split clause 116,176]
181. 4 <=> v3_ordinal1(sK1) [avatar definition]
183. v3_ordinal1(sK1) <- (4) [avatar component clause 181] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
184. 4 [avatar split clause 115,181]
186. 5 <=> v3_ordinal1(sK0) [avatar definition]
188. v3_ordinal1(sK0) <- (5) [avatar component clause 186] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
189. 5 [avatar split clause 114,186]
196. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
198. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 196] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
199. 7 [avatar split clause 120,196]
382. v3_ordinal1(sK2(X0)) | ~v3_ordinal1(X0) [resolution 133,145] {a:1,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:3,thDist:-3}
416. r2_hidden(sK2(X0),X0) | v1_xboole_0(X0) [resolution 136,145] {a:1,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:3,thDist:-3}
445. r2_hidden(k11_ordinal2(X0,sK1),k1_xboole_0) | k1_xboole_0 = sK1 | ~r2_hidden(X0,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK1) <- (3) [superposition 132,178] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
449. r2_hidden(k11_ordinal2(X0,sK1),k1_xboole_0) | ~r2_hidden(X0,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK1) <- (~1, 3) [subsumption resolution 445,168] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
450. r2_hidden(k11_ordinal2(X0,sK1),k1_xboole_0) | ~r2_hidden(X0,sK0) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK1) <- (~1, 3, 5) [subsumption resolution 449,188] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
451. r2_hidden(k11_ordinal2(X0,sK1),k1_xboole_0) | ~r2_hidden(X0,sK0) | ~v3_ordinal1(X0) <- (~1, 3, 4, 5) [subsumption resolution 450,183] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
466. ~r2_hidden(X2,sK0) | ~v3_ordinal1(X2) | ~v1_xboole_0(k1_xboole_0) <- (~1, 3, 4, 5) [resolution 451,142] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
469. ~r2_hidden(X2,sK0) | ~v3_ordinal1(X2) <- (~1, 3, 4, 5, 7) [subsumption resolution 466,198] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
475. 45 <=> ! [X1] : (~r2_hidden(X1,sK0) | ~v3_ordinal1(X1)) [avatar definition]
476. ~r2_hidden(X1,sK0) | ~v3_ordinal1(X1) <- (45) [avatar component clause 475] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
484. 45 | 1 | ~3 | ~4 | ~5 | ~7 [avatar split clause 469,196,186,181,176,166,475]
525. v1_xboole_0(sK0) | ~v3_ordinal1(sK2(sK0)) <- (45) [resolution 416,476] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
527. 46 <=> v3_ordinal1(sK2(sK0)) [avatar definition]
529. ~v3_ordinal1(sK2(sK0)) <- (~46) [avatar component clause 527] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
531. 47 <=> v1_xboole_0(sK0) [avatar definition]
533. v1_xboole_0(sK0) <- (47) [avatar component clause 531] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
534. ~46 | 47 | ~45 [avatar split clause 525,475,531,527]
537. ~v3_ordinal1(sK0) <- (~46) [resolution 529,382] {a:4,w:2,wCS:200,goal:1,sine:3,thAx:0,allAx:11,thDist:-11}
540. $false <- (5, ~46) [subsumption resolution 537,188] {a:4,w:0,goal:1,sine:3,thAx:0,allAx:11,thDist:-11}
541. ~5 | 46 [avatar contradiction clause 540]
549. $false <- (~2, 7, 47) [unit resulting resolution 198,173,533,141] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
555. 2 | ~7 | ~47 [avatar contradiction clause 549]
556. $false [avatar sat refutation 169,174,179,184,189,199,484,534,541,555] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:329,thDist:-329}
% SZS output end Proof for ordinal3__t31_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 44
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 51
% Generated clauses: 259
% Active clauses: 99
% Passive clauses: 140
% Final active clauses: 85
% Final passive clauses: 36
% Main loop iterations started: 104
% 
% >>> Simplifying Inferences
% Duplicate literals: 7
% Fw subsumption resolutions: 14
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Equational tautologies: 2
% Forward subsumptions: 41
% Backward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 95
% Unit resulting resolution: 6
% Factoring: 1
% Forward superposition: 18
% Backward superposition: 9
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 58
% SAT solver unit clauses: 18
% SAT solver binary clauses: 32
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% forward subsumption: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

