i: [1,37,0,2,1,1,0,0] t31_orders_1
a: [3,62,0,2,1,37]
s: 62
i: [1,36,0,2,1,1,0,0] t31_orders_1
a: [3,67,0,2,1,36]
s: 67
i: [1,38,0,4,1,1,0,0] t31_orders_1
a: [3,57,0,4,1,38]
s: 57
i: [1,50,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 50
i: [1,42,0,6,3,0,0,1] d3_orders_1
s: 42
i: [1,43,0,6,3,0,0,1] d3_orders_1
s: 43
d: [2,69,1,4,2,1,40,42,62]
d: [2,70,1,2,1,2,21,69,67]
a: [3,74,1,2,1,70]
s: 74
i: [1,51,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 51
d: [2,76,1,4,2,1,40,43,62]
d: [2,77,1,2,1,2,21,76,67]
a: [3,81,1,2,1,77]
s: 81
i: [1,40,0,8,3,0,0,2] d16_relat_2
s: 40
i: [1,41,0,8,3,0,0,2] d16_relat_2
s: 41
i: [1,44,0,8,4,0,0,1] d3_orders_1
s: 44
i: [1,45,0,8,3,0,0,2] d9_relat_2
s: 45
i: [1,46,0,8,3,0,0,2] d9_relat_2
s: 46
i: [1,47,0,8,3,0,0,1] d6_orders_1
s: 47
i: [1,48,0,8,3,0,0,1] d6_orders_1
s: 48
i: [1,53,0,8,3,0,0,3] fc3_relat_1
s: 53
i: [1,39,0,10,2,0,0,255] d6_relat_1
s: 39
i: [1,49,0,11,4,0,0,1] d6_orders_1
s: 49
i: [1,52,0,12,5,0,0,2] fc4_relat_2
s: 52
d: [2,92,1,11,4,0,44,53,39]
s: 92
d: [2,93,1,14,5,0,40,49,45]
d: [2,94,1,12,4,0,18,93]
s: 94
d: [2,95,1,20,7,0,40,52,44]
d: [2,101,1,16,6,0,21,95,53]
s: 101
d: [2,102,2,12,5,0,40,94,40]
d: [2,103,2,10,4,0,18,102]
s: 103
d: [2,109,3,0,0,4,52,81,74,67,57,103]
e: 109
d: [2,110,3,6,3,1,40,103,57]
d: [2,116,3,4,2,2,21,110,67]
d: [2,117,3,2,1,3,21,116,74]
d: [2,118,3,0,0,4,21,117,81]
e: 118
f: 118,77,70,36,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t31_orders_1
% SZS output start Proof for orders_1__t31_orders_1
1. ! [X0] : (v1_relat_1(X0) => (v1_orders_1(X0) => r1_orders_1(X0,k1_relat_1(X0)))) [input t31_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => (v1_orders_1(X0) => r1_orders_1(X0,k1_relat_1(X0)))) [negated conjecture 1]
4. ! [X0] : (v1_relat_1(X0) => (v8_relat_2(X0) <=> r8_relat_2(X0,k1_relat_1(X0)))) [input d16_relat_2]
5. ! [X0] : (v1_relat_1(X0) => (v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0)))) [input d3_orders_1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_orders_1(X0,X1) <=> (r8_relat_2(X0,X1) & r1_relat_2(X0,X1)))) [input d6_orders_1]
8. ! [X0] : (v1_relat_1(X0) => (v1_relat_2(X0) <=> r1_relat_2(X0,k1_relat_1(X0)))) [input d9_relat_2]
17. ? [X0] : ((~r1_orders_1(X0,k1_relat_1(X0)) & v1_orders_1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
18. ? [X0] : (~r1_orders_1(X0,k1_relat_1(X0)) & v1_orders_1(X0) & v1_relat_1(X0)) [flattening 17]
20. ! [X0] : ((v8_relat_2(X0) <=> r8_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
21. ! [X0] : ((v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 5]
22. ! [X0] : ((v1_relat_2(X0) <=> r1_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 8]
23. ! [X0] : (! [X1] : (r1_orders_1(X0,X1) <=> (r8_relat_2(X0,X1) & r1_relat_2(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 6]
28. ? [X0] : (~r1_orders_1(X0,k1_relat_1(X0)) & v1_orders_1(X0) & v1_relat_1(X0)) => (~r1_orders_1(sK0,k1_relat_1(sK0)) & v1_orders_1(sK0) & v1_relat_1(sK0)) [choice axiom]
29. ~r1_orders_1(sK0,k1_relat_1(sK0)) & v1_orders_1(sK0) & v1_relat_1(sK0) [skolemisation 18,28]
30. ! [X0] : (((v8_relat_2(X0) | ~r8_relat_2(X0,k1_relat_1(X0))) & (r8_relat_2(X0,k1_relat_1(X0)) | ~v8_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 20]
31. ! [X0] : (((v1_orders_1(X0) | (~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 21]
32. ! [X0] : (((v1_orders_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 31]
33. ! [X0] : (((v1_relat_2(X0) | ~r1_relat_2(X0,k1_relat_1(X0))) & (r1_relat_2(X0,k1_relat_1(X0)) | ~v1_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 22]
34. ! [X0] : (! [X1] : ((r1_orders_1(X0,X1) | (~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1))) & ((r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r1_orders_1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 23]
35. ! [X0] : (! [X1] : ((r1_orders_1(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1)) & ((r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r1_orders_1(X0,X1))) | ~v1_relat_1(X0)) [flattening 34]
36. v1_relat_1(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. v1_orders_1(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. ~r1_orders_1(sK0,k1_relat_1(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. r8_relat_2(X0,k1_relat_1(X0)) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
42. ~v1_orders_1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 32] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. ~v1_orders_1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 32] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. r1_relat_2(X0,k1_relat_1(X0)) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
49. ~r1_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | r1_orders_1(X0,X1) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. 1 <=> r1_orders_1(sK0,k1_relat_1(sK0)) [avatar definition]
57. ~r1_orders_1(sK0,k1_relat_1(sK0)) <- (~1) [avatar component clause 55] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
58. ~1 [avatar split clause 38,55]
60. 2 <=> v1_orders_1(sK0) [avatar definition]
62. v1_orders_1(sK0) <- (2) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. 2 [avatar split clause 37,60]
65. 3 <=> v1_relat_1(sK0) [avatar definition]
67. v1_relat_1(sK0) <- (3) [avatar component clause 65] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. 3 [avatar split clause 36,65]
69. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 42,62] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
70. v1_relat_2(sK0) <- (2, 3) [subsumption resolution 69,67] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
72. 4 <=> v1_relat_2(sK0) [avatar definition]
74. v1_relat_2(sK0) <- (4) [avatar component clause 72] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
75. 4 | ~2 | ~3 [avatar split clause 70,65,60,72]
76. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 43,62] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
77. v8_relat_2(sK0) <- (2, 3) [subsumption resolution 76,67] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
79. 5 <=> v8_relat_2(sK0) [avatar definition]
81. v8_relat_2(sK0) <- (5) [avatar component clause 79] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. 5 | ~2 | ~3 [avatar split clause 77,65,60,79]
93. ~r8_relat_2(X0,k1_relat_1(X0)) | r1_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [resolution 49,45] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
94. ~r8_relat_2(X0,k1_relat_1(X0)) | r1_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) [duplicate literal removal 93] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
102. r1_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [resolution 94,40] {a:2,w:12,wCS:4800,sine:1,thAx:0,allAx:3,thDist:-3}
103. r1_orders_1(X0,k1_relat_1(X0)) | ~v1_relat_1(X0) | ~v1_relat_2(X0) | ~v8_relat_2(X0) [duplicate literal removal 102] {a:2,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
110. ~v1_relat_1(sK0) | ~v1_relat_2(sK0) | ~v8_relat_2(sK0) <- (~1) [resolution 103,57] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
116. ~v1_relat_2(sK0) | ~v8_relat_2(sK0) <- (~1, 3) [subsumption resolution 110,67] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
117. ~v8_relat_2(sK0) <- (~1, 3, 4) [subsumption resolution 116,74] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
118. $false <- (~1, 3, 4, 5) [subsumption resolution 117,81] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
119. 1 | ~3 | ~4 | ~5 [avatar contradiction clause 118]
120. $false [avatar sat refutation 58,63,68,75,82,119] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:39,thDist:-39}
% SZS output end Proof for orders_1__t31_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 1
% 
% >>> Saturation
% Initial clauses: 18
% Generated clauses: 62
% Active clauses: 24
% Passive clauses: 26
% Final active clauses: 24
% Final passive clauses: 2
% Main loop iterations started: 24
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 1
% Forward superposition: 11
% Backward superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward superposition index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

