i: [1,51,0,2,1,1,0,0] t185_relat_1
a: [3,91,0,2,1,51]
s: 91
i: [1,53,0,3,1,1,0,0] t185_relat_1
a: [3,81,0,3,1,53]
s: 81
i: [1,52,0,3,1,1,0,0] t185_relat_1
a: [3,86,0,3,1,52]
s: 86
i: [1,54,0,5,1,1,0,0] t185_relat_1
a: [3,76,0,5,1,54]
s: 76
i: [1,66,0,3,1,0,0,1] reflexivity_r1_tarski
s: 66
i: [1,65,0,4,1,0,0,2] existence_m1_subset_1
s: 65
i: [1,67,0,4,1,0,0,1] fc6_relat_1
s: 67
i: [1,70,0,4,1,0,0,1] rc3_relat_1
s: 70
i: [1,71,0,5,1,0,0,1] rc3_relat_1
s: 71
i: [1,72,0,5,1,0,0,1] rc3_relat_1
s: 72
i: [1,61,0,7,2,0,0,2] t3_subset
s: 61
d: [2,93,1,5,1,0,40,61,65]
s: 93
i: [1,62,0,7,2,0,0,2] t3_subset
s: 62
i: [1,56,0,8,3,0,0,2] cc2_relat_1
s: 56
i: [1,55,0,9,2,0,0,255] t7_relat_1
s: 55
i: [1,57,0,9,3,0,0,255] d18_relat_1
s: 57
d: [2,95,1,6,2,0,40,56,65]
s: 95
i: [1,58,0,9,3,0,0,255] d18_relat_1
s: 58
d: [2,97,1,6,2,0,40,58,66]
s: 97
i: [1,59,0,9,3,0,0,255] d19_relat_1
s: 59
i: [1,60,0,9,3,0,0,255] d19_relat_1
s: 60
d: [2,101,1,6,2,0,40,60,66]
s: 101
i: [1,63,0,9,3,0,0,1] t1_xboole_1
s: 63
d: [2,96,1,7,3,0,40,56,62]
s: 96
i: [1,68,0,12,4,0,0,1] cc5_relat_1
s: 68
i: [1,69,0,12,4,0,0,1] cc6_relat_1
s: 69
d: [2,106,2,8,2,0,40,63,93]
s: 106
i: [1,64,0,13,3,0,0,1] t96_zfmisc_1
s: 64
d: [2,120,3,7,1,0,40,106,93]
s: 120
d: [2,115,1,10,3,0,40,68,65]
s: 115
d: [2,117,1,10,3,0,40,69,65]
s: 117
d: [2,125,4,8,2,0,40,120,96]
s: 125
d: [2,116,1,11,4,0,40,68,62]
s: 116
d: [2,127,4,9,1,0,40,120,106]
s: 127
d: [2,118,1,11,4,0,40,69,62]
s: 118
d: [2,107,1,12,3,0,40,63,55]
s: 107
d: [2,113,2,10,4,0,40,96,57]
s: 113
d: [2,108,1,12,4,0,40,63,57]
s: 108
d: [2,160,3,7,3,1,40,113,86]
d: [2,164,3,5,2,2,21,160,91]
a: [3,168,3,2,1,164]
s: 168
d: [2,109,1,12,4,0,40,63,59]
s: 109
d: [2,98,1,14,3,0,40,58,55]
s: 98
d: [2,162,3,11,3,0,40,113,71]
d: [2,175,3,7,2,0,21,162,70]
s: 175
d: [2,102,1,14,3,0,40,60,55]
s: 102
d: [2,114,2,10,4,0,40,96,59]
s: 114
d: [2,124,1,16,4,0,40,64,63]
s: 124
d: [2,177,2,11,3,0,40,108,93]
s: 177
d: [2,192,3,7,3,1,40,114,81]
d: [2,197,3,5,2,2,21,192,91]
a: [3,201,3,2,1,197]
s: 201
d: [2,183,2,11,3,0,40,109,93]
s: 183
d: [2,194,3,11,3,0,40,114,72]
d: [2,208,3,7,2,0,21,194,70]
s: 208
d: [2,131,2,18,4,0,40,116,55]
d: [2,135,2,12,3,0,21,131,67]
s: 135
d: [2,145,2,18,4,0,40,118,55]
d: [2,149,2,12,3,0,21,145,67]
s: 149
d: [2,234,3,16,3,0,40,135,97]
d: [2,236,3,10,2,0,21,234,67]
s: 236
d: [2,133,2,14,5,0,40,116,57]
s: 133
d: [2,238,3,16,3,0,40,149,101]
d: [2,240,3,10,2,0,21,238,67]
s: 240
d: [2,134,2,14,5,0,40,116,59]
s: 134
d: [2,147,2,14,5,0,40,118,57]
s: 147
d: [2,126,4,10,2,0,40,120,63]
s: 126
d: [2,148,2,14,5,0,40,118,59]
s: 148
d: [2,138,5,10,2,0,40,127,96]
s: 138
d: [2,154,2,15,4,0,40,107,63]
s: 154
d: [2,155,2,15,4,0,40,107,58]
s: 155
d: [2,269,3,16,5,0,40,154,96]
d: [2,275,3,10,4,0,21,269,67]
s: 275
d: [2,156,2,15,4,0,40,107,60]
s: 156
d: [2,281,4,9,3,0,40,275,93]
s: 281
d: [2,151,2,21,5,0,40,107,118]
d: [2,157,2,15,4,0,21,151,67]
s: 157
d: [2,152,2,21,5,0,40,107,116]
d: [2,158,2,15,4,0,21,152,67]
s: 158
d: [2,121,3,11,3,0,40,106,57]
s: 121
d: [2,180,2,15,5,0,40,108,57]
s: 180
d: [2,315,4,14,2,0,40,121,71]
d: [2,319,4,8,1,0,21,315,70]
s: 319
d: [2,181,2,15,5,0,40,108,59]
s: 181
d: [2,186,2,15,5,0,40,109,57]
s: 186
d: [2,333,5,13,2,0,40,319,58]
d: [2,344,5,7,1,0,21,333,70]
s: 344
d: [2,187,2,15,5,0,40,109,59]
s: 187
d: [2,337,5,9,2,0,40,319,96]
s: 337
d: [2,209,2,15,4,0,40,124,55]
s: 209
d: [2,213,2,15,3,0,40,124,93]
s: 213
d: [2,390,3,10,3,1,40,209,76]
d: [2,393,3,8,2,2,21,390,91]
a: [3,397,3,4,1,393]
s: 397
d: [2,412,4,0,0,3,52,91,81,397,59]
e: 412
d: [2,413,4,0,0,3,52,91,81,66,397,109]
e: 413
d: [2,414,4,5,2,1,40,397,59]
d: [2,417,4,2,1,2,21,414,81]
d: [2,418,4,0,0,3,21,417,91]
e: 418
r: 397
d: [2,132,2,20,5,0,40,116,64]
d: [2,136,2,16,4,0,21,132,67]
s: 136
a: [3,396,3,4,1,393]
s: 396
d: [2,146,2,20,5,0,40,118,64]
d: [2,150,2,16,4,0,21,146,67]
s: 150
d: [2,210,2,18,5,0,40,124,107]
s: 210
a: [3,205,3,3,1,197]
s: 205
d: [2,211,2,19,5,0,40,124,64]
s: 211
a: [3,401,3,4,1,393]
s: 401
d: [2,465,4,0,0,3,52,91,86,401,57]
e: 465
d: [2,466,4,0,0,3,52,91,86,66,401,108]
e: 466
d: [2,467,4,5,2,1,40,401,57]
d: [2,470,4,2,1,2,21,467,86]
d: [2,471,4,0,0,3,21,470,91]
e: 471
f: 471,418,393,51,52,53,54
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t185_relat_1
% SZS output start Proof for relat_1__t185_relat_1
1. ! [X0] : ! [X1] : ! [X2] : ((v5_relat_1(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r1_tarski(X2,k2_zfmisc_1(X0,X1))) [input t185_relat_1]
2. ~! [X0] : ! [X1] : ! [X2] : ((v5_relat_1(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r1_tarski(X2,k2_zfmisc_1(X0,X1))) [negated conjecture 1]
6. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0))) [input d18_relat_1]
7. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v5_relat_1(X1,X0) <=> r1_tarski(k10_xtuple_0(X1),X0))) [input d19_relat_1]
17. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
19. ! [X0] : (v1_relat_1(X0) => r1_tarski(X0,k2_zfmisc_1(k9_xtuple_0(X0),k10_xtuple_0(X0)))) [input t7_relat_1]
20. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3))) [input t96_zfmisc_1]
21. ~! [X0,X1,X2] : ((v5_relat_1(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r1_tarski(X2,k2_zfmisc_1(X0,X1))) [flattening 2]
22. ! [X0,X1] : (v1_relat_1(X1) => (v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0))) [flattening 6]
23. ! [X0,X1] : (v1_relat_1(X1) => (v5_relat_1(X1,X0) <=> r1_tarski(k10_xtuple_0(X1),X0))) [flattening 7]
25. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 17]
26. ! [X0,X1,X2,X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3))) [flattening 20]
28. ? [X0,X1,X2] : (~r1_tarski(X2,k2_zfmisc_1(X0,X1)) & (v5_relat_1(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2))) [ennf transformation 21]
29. ? [X0,X1,X2] : (~r1_tarski(X2,k2_zfmisc_1(X0,X1)) & v5_relat_1(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2)) [flattening 28]
30. ! [X0] : (r1_tarski(X0,k2_zfmisc_1(k9_xtuple_0(X0),k10_xtuple_0(X0))) | ~v1_relat_1(X0)) [ennf transformation 19]
32. ! [X0,X1] : ((v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0)) | ~v1_relat_1(X1)) [ennf transformation 22]
33. ! [X0,X1] : ((v5_relat_1(X1,X0) <=> r1_tarski(k10_xtuple_0(X1),X0)) | ~v1_relat_1(X1)) [ennf transformation 23]
34. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 25]
35. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 34]
36. ! [X0,X1,X2,X3] : (r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | (~r1_tarski(X2,X3) | ~r1_tarski(X0,X1))) [ennf transformation 26]
37. ! [X0,X1,X2,X3] : (r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | ~r1_tarski(X2,X3) | ~r1_tarski(X0,X1)) [flattening 36]
42. ? [X0,X1,X2] : (~r1_tarski(X2,k2_zfmisc_1(X0,X1)) & v5_relat_1(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => (~r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) & v5_relat_1(sK2,sK1) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
43. ~r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) & v5_relat_1(sK2,sK1) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2) [skolemisation 29,42]
44. ! [X0,X1] : (((v4_relat_1(X1,X0) | ~r1_tarski(k9_xtuple_0(X1),X0)) & (r1_tarski(k9_xtuple_0(X1),X0) | ~v4_relat_1(X1,X0))) | ~v1_relat_1(X1)) [nnf transformation 32]
45. ! [X0,X1] : (((v5_relat_1(X1,X0) | ~r1_tarski(k10_xtuple_0(X1),X0)) & (r1_tarski(k10_xtuple_0(X1),X0) | ~v5_relat_1(X1,X0))) | ~v1_relat_1(X1)) [nnf transformation 33]
51. v1_relat_1(sK2) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. v4_relat_1(sK2,sK0) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. v5_relat_1(sK2,sK1) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. ~r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) [cnf transformation 43] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. r1_tarski(X0,k2_zfmisc_1(k9_xtuple_0(X0),k10_xtuple_0(X0))) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
57. r1_tarski(k9_xtuple_0(X1),X0) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 44] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
59. r1_tarski(k10_xtuple_0(X1),X0) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 45] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
63. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 35] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | ~r1_tarski(X2,X3) | ~r1_tarski(X0,X1) [cnf transformation 37] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
74. 1 <=> r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) [avatar definition]
76. ~r1_tarski(sK2,k2_zfmisc_1(sK0,sK1)) <- (~1) [avatar component clause 74] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
77. ~1 [avatar split clause 54,74]
79. 2 <=> v5_relat_1(sK2,sK1) [avatar definition]
81. v5_relat_1(sK2,sK1) <- (2) [avatar component clause 79] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
82. 2 [avatar split clause 53,79]
84. 3 <=> v4_relat_1(sK2,sK0) [avatar definition]
86. v4_relat_1(sK2,sK0) <- (3) [avatar component clause 84] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. 3 [avatar split clause 52,84]
89. 4 <=> v1_relat_1(sK2) [avatar definition]
91. v1_relat_1(sK2) <- (4) [avatar component clause 89] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
92. 4 [avatar split clause 51,89]
124. ~r1_tarski(X8,k2_zfmisc_1(X6,X4)) | ~r1_tarski(X6,X7) | r1_tarski(X8,k2_zfmisc_1(X7,X5)) | ~r1_tarski(X4,X5) [resolution 64,63] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
209. r1_tarski(X0,k2_zfmisc_1(X1,X2)) | ~r1_tarski(k9_xtuple_0(X0),X1) | ~r1_tarski(k10_xtuple_0(X0),X2) | ~v1_relat_1(X0) [resolution 124,55] {a:2,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
390. ~r1_tarski(k9_xtuple_0(sK2),sK0) | ~r1_tarski(k10_xtuple_0(sK2),sK1) | ~v1_relat_1(sK2) <- (~1) [resolution 209,76] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
393. ~r1_tarski(k9_xtuple_0(sK2),sK0) | ~r1_tarski(k10_xtuple_0(sK2),sK1) <- (~1, 4) [subsumption resolution 390,91] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
395. 9 <=> r1_tarski(k10_xtuple_0(sK2),sK1) [avatar definition]
397. ~r1_tarski(k10_xtuple_0(sK2),sK1) <- (~9) [avatar component clause 395] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
399. 10 <=> r1_tarski(k9_xtuple_0(sK2),sK0) [avatar definition]
401. ~r1_tarski(k9_xtuple_0(sK2),sK0) <- (~10) [avatar component clause 399] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
402. ~9 | ~10 | 1 | ~4 [avatar split clause 393,89,74,399,395]
414. ~v5_relat_1(sK2,sK1) | ~v1_relat_1(sK2) <- (~9) [resolution 397,59] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
417. ~v1_relat_1(sK2) <- (2, ~9) [subsumption resolution 414,81] {a:4,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
418. $false <- (2, 4, ~9) [subsumption resolution 417,91] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
419. ~2 | ~4 | 9 [avatar contradiction clause 418]
467. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~10) [resolution 401,57] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
470. ~v1_relat_1(sK2) <- (3, ~10) [subsumption resolution 467,86] {a:4,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
471. $false <- (3, 4, ~10) [subsumption resolution 470,91] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
472. ~3 | ~4 | 10 [avatar contradiction clause 471]
473. $false [avatar sat refutation 77,82,87,92,402,419,472] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:108,thDist:-108}
% SZS output end Proof for relat_1__t185_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 19
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 22
% Generated clauses: 375
% Active clauses: 84
% Passive clauses: 221
% Final active clauses: 83
% Final passive clauses: 137
% Main loop iterations started: 84
% 
% >>> Simplifying Inferences
% Duplicate literals: 16
% Fw subsumption resolutions: 84
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 27
% 
% >>> Generating Inferences
% Binary resolution: 241
% Unit resulting resolution: 4
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 4
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% forward subsumption: 0.002 s
% other: 0.004 s ( own 0.001 s ) 

