i: [1,220,0,2,1,1,0,0] t58_xreal_1
a: [3,358,0,2,1,220]
s: 358
i: [1,224,0,7,2,1,0,0] t58_xreal_1
a: [3,352,0,3,1,224]
s: 352
a: [3,347,0,4,1,224]
s: 347
i: [1,225,0,2,1,0,0,2] fc1_numbers
a: [3,363,0,2,1,225]
s: 363
i: [1,226,0,2,1,0,0,3] spc1_boole
a: [3,368,0,2,1,226]
s: 368
i: [1,227,0,2,1,0,0,2] fc1_xboole_0
a: [3,373,0,2,1,227]
s: 373
i: [1,240,0,2,1,0,0,3] fc6_ordinal1
a: [3,428,0,2,1,240]
s: 428
i: [1,239,0,2,1,0,0,3] fc6_ordinal1
a: [3,433,0,2,1,239]
s: 433
i: [1,241,0,2,1,0,0,2] spc1_numerals
a: [3,453,0,2,1,241]
s: 453
i: [1,310,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,457,0,2,1,310]
s: 457
i: [1,312,0,2,1,0,0,2] rc2_xboole_0
a: [3,465,0,2,1,312]
s: 465
i: [1,313,0,2,1,0,0,4] rc4_ordinal1
a: [3,470,0,2,1,313]
s: 470
i: [1,314,0,2,1,0,0,4] rc1_ordinal1
a: [3,475,0,2,1,314]
s: 475
i: [1,315,0,2,1,0,0,1] rc1_xreal_0
a: [3,480,0,2,1,315]
s: 480
i: [1,316,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,485,0,2,1,316]
s: 485
i: [1,317,0,2,1,0,0,2] rc1_xxreal_0
a: [3,490,0,2,1,317]
s: 490
d: [2,627,0,0,0,2,21,490,457]
e: 627
r: 457
s: 490
i: [1,318,0,2,1,0,0,2] rc1_xboole_0
a: [3,495,0,2,1,318]
s: 495
i: [1,322,0,2,1,0,0,255] rc2_ordinal1
a: [3,500,0,2,1,322]
s: 500
i: [1,321,0,2,1,0,0,255] rc2_ordinal1
a: [3,505,0,2,1,321]
s: 505
i: [1,320,0,2,1,0,0,255] rc2_ordinal1
a: [3,510,0,2,1,320]
s: 510
i: [1,319,0,2,1,0,0,255] rc2_ordinal1
a: [3,515,0,2,1,319]
s: 515
i: [1,324,0,2,1,0,0,4] rc5_ordinal1
a: [3,520,0,2,1,324]
s: 520
i: [1,323,0,2,1,0,0,4] rc5_ordinal1
a: [3,525,0,2,1,323]
s: 525
i: [1,326,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,530,0,2,1,326]
s: 530
i: [1,325,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,535,0,2,1,325]
s: 535
i: [1,330,0,2,1,0,0,2] rc3_xreal_0
a: [3,540,0,2,1,330]
s: 540
i: [1,329,0,2,1,0,0,2] rc3_xreal_0
a: [3,545,0,2,1,329]
s: 545
i: [1,328,0,2,1,0,0,2] rc3_xreal_0
a: [3,550,0,2,1,328]
s: 550
i: [1,327,0,2,1,0,0,2] rc3_xreal_0
a: [3,555,0,2,1,327]
s: 555
i: [1,334,0,2,1,0,0,2] rc2_xreal_0
a: [3,560,0,2,1,334]
s: 560
i: [1,333,0,2,1,0,0,2] rc2_xreal_0
a: [3,565,0,2,1,333]
s: 565
i: [1,332,0,2,1,0,0,2] rc2_xreal_0
a: [3,570,0,2,1,332]
s: 570
i: [1,331,0,2,1,0,0,2] rc2_xreal_0
a: [3,575,0,2,1,331]
s: 575
i: [1,336,0,2,1,0,0,2] rc3_xxreal_0
a: [3,580,0,2,1,336]
s: 580
i: [1,335,0,2,1,0,0,2] rc3_xxreal_0
a: [3,585,0,2,1,335]
s: 585
i: [1,338,0,2,1,0,0,2] rc2_xxreal_0
a: [3,590,0,2,1,338]
s: 590
i: [1,337,0,2,1,0,0,2] rc2_xxreal_0
a: [3,595,0,2,1,337]
s: 595
i: [1,342,0,2,1,0,0,2] rc4_xreal_0
a: [3,600,0,2,1,342]
s: 600
i: [1,341,0,2,1,0,0,2] rc4_xreal_0
a: [3,605,0,2,1,341]
s: 605
i: [1,340,0,2,1,0,0,2] rc4_xreal_0
a: [3,610,0,2,1,340]
s: 610
i: [1,339,0,2,1,0,0,2] rc4_xreal_0
a: [3,615,0,2,1,339]
s: 615
i: [1,344,0,2,1,0,0,2] rc4_xxreal_0
a: [3,620,0,2,1,344]
s: 620
i: [1,343,0,2,1,0,0,2] rc4_xxreal_0
a: [3,625,0,2,1,343]
s: 625
i: [1,304,0,3,1,0,0,255] reflexivity_r1_tarski
s: 304
i: [1,228,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,378,0,3,1,228]
s: 378
i: [1,229,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,383,0,3,1,229]
s: 383
r: 373
d: [2,629,0,2,1,2,23,373,383]
a: [3,633,0,2,1,629]
s: 633
i: [1,230,0,3,1,0,0,2] t1_numerals
a: [3,388,0,3,1,230]
s: 388
i: [1,231,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,393,0,3,1,231]
s: 393
i: [1,244,0,3,1,0,0,2] spc1_numerals
a: [3,438,0,3,1,244]
s: 438
i: [1,243,0,3,1,0,0,2] spc1_numerals
a: [3,443,0,3,1,243]
s: 443
d: [2,635,0,3,1,2,22,388,383]
a: [3,639,0,3,1,635]
s: 639
d: [2,641,0,3,1,2,22,443,378]
a: [3,645,0,3,1,641]
s: 645
i: [1,246,0,4,2,0,0,4] cc6_ordinal1
s: 246
i: [1,247,0,4,2,0,0,4] cc2_xreal_0
s: 247
d: [2,647,1,2,1,1,40,246,470]
a: [3,652,1,2,1,647]
s: 652
i: [1,248,0,4,2,0,0,4] cc2_xcmplx_0
s: 248
d: [2,648,1,2,1,1,40,246,520]
a: [3,657,1,2,1,648]
s: 657
i: [1,249,0,4,2,0,0,4] cc2_xxreal_0
s: 249
i: [1,250,0,4,2,0,0,255] cc1_ordinal1
s: 250
d: [2,659,1,2,1,1,40,247,470]
a: [3,664,1,2,1,659]
s: 664
i: [1,251,0,4,2,0,0,255] cc1_ordinal1
s: 251
d: [2,660,1,2,1,1,40,247,520]
a: [3,669,1,2,1,660]
s: 669
i: [1,253,0,4,2,0,0,2] cc3_xreal_0
s: 253
i: [1,254,0,4,2,0,0,2] cc4_xreal_0
s: 254
d: [2,745,1,2,1,1,40,253,358]
a: [3,755,1,2,1,745]
s: 755
i: [1,267,0,4,2,0,0,255] cc9_ordinal1
s: 267
d: [2,764,1,2,1,1,40,254,358]
a: [3,774,1,2,1,764]
s: 774
i: [1,268,0,4,2,0,0,4] cc7_ordinal1
s: 268
i: [1,269,0,4,2,0,0,4] cc3_ordinal1
s: 269
d: [2,671,1,2,1,1,40,248,470]
a: [3,676,1,2,1,671]
s: 676
i: [1,303,0,4,1,0,0,3] existence_m1_subset_1
s: 303
d: [2,672,1,2,1,1,40,248,520]
a: [3,681,1,2,1,672]
s: 681
i: [1,232,0,4,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r1_rm1
a: [3,398,0,4,1,232]
s: 398
i: [1,233,0,4,1,0,0,2] dt_k5_numbers
a: [3,403,0,4,1,233]
s: 403
d: [2,683,1,2,1,1,40,249,470]
a: [3,688,1,2,1,683]
s: 688
i: [1,234,0,4,1,0,0,3] rqLessOrEqual__r1_xxreal_0__rm1_r1
a: [3,408,0,4,1,234]
s: 408
d: [2,684,1,2,1,1,40,249,520]
a: [3,693,1,2,1,684]
s: 693
i: [1,235,0,4,1,0,0,1] dt_k6_numbers
a: [3,413,0,4,1,235]
s: 413
i: [1,242,0,4,1,0,0,2] spc1_numerals
a: [3,448,0,4,1,242]
s: 448
d: [2,695,1,2,1,1,40,250,428]
a: [3,703,1,2,1,695]
s: 703
d: [2,855,0,4,1,2,22,403,378]
a: [3,859,0,4,1,855]
s: 859
d: [2,697,1,2,1,1,40,250,475]
a: [3,713,1,2,1,697]
s: 713
d: [2,861,0,4,1,2,22,413,378]
a: [3,865,0,4,1,861]
s: 865
d: [2,867,0,4,1,2,22,448,378]
a: [3,871,0,4,1,867]
s: 871
d: [2,720,1,2,1,1,40,251,428]
a: [3,728,1,2,1,720]
s: 728
i: [1,255,0,5,2,0,0,1] fc3_xreal_0
s: 255
d: [2,722,1,2,1,1,40,251,475]
a: [3,738,1,2,1,722]
s: 738
i: [1,256,0,5,2,0,0,1] fc3_xreal_0
s: 256
i: [1,265,0,5,2,0,0,1] dt_k4_xcmplx_0
s: 265
d: [2,747,1,2,1,1,40,253,480]
a: [3,761,1,2,1,747]
s: 761
i: [1,270,0,5,2,0,0,2] t6_boole
s: 270
d: [2,766,1,2,1,1,40,254,480]
a: [3,780,1,2,1,766]
s: 780
i: [1,271,0,5,2,0,0,3] cc8_ordinal1
s: 271
i: [1,272,0,5,2,0,0,2] cc1_xreal_0
s: 272
d: [2,783,1,2,1,1,40,267,633]
a: [3,790,1,2,1,783]
s: 790
i: [1,273,0,5,2,0,0,2] cc1_xcmplx_0
s: 273
d: [2,784,1,2,1,1,40,267,495]
a: [3,795,1,2,1,784]
s: 795
i: [1,300,0,5,2,0,0,255] t7_boole
s: 300
i: [1,236,0,5,1,0,0,2] rqRealNeg__k4_xcmplx_0__rm1_r1
a: [3,418,0,5,1,236]
s: 418
d: [2,785,1,2,1,1,40,267,615]
a: [3,800,1,2,1,785]
s: 800
i: [1,238,0,5,1,0,0,3] rqLessOrEqual__r1_xxreal_0__rm1_rm1
a: [3,423,0,5,1,238]
s: 423
d: [2,786,1,2,1,1,40,267,625]
a: [3,805,1,2,1,786]
s: 805
a: [3,460,0,5,2,310]
s: 460
d: [2,875,0,5,2,1,22,270,383]
s: 875
d: [2,930,1,5,2,2,44,352,875]
s: 930
i: [1,284,0,6,3,0,0,255] cc2_ordinal1
s: 284
d: [2,929,1,6,2,2,44,347,875]
s: 929
i: [1,287,0,6,3,0,0,2] cc5_xxreal_0
s: 287
i: [1,289,0,6,3,0,0,2] cc5_xxreal_0
s: 289
d: [2,938,2,6,2,3,40,929,460]
a: [3,947,2,3,1,938]
s: 947
i: [1,290,0,6,3,0,0,2] cc3_xxreal_0
s: 290
d: [2,807,1,2,1,1,40,268,633]
a: [3,814,1,2,1,807]
s: 814
i: [1,292,0,6,3,0,0,2] cc3_xxreal_0
s: 292
i: [1,294,0,6,3,0,0,2] cc7_xxreal_0
s: 294
d: [2,808,1,2,1,1,40,268,495]
a: [3,819,1,2,1,808]
s: 819
i: [1,295,0,6,3,0,0,2] cc7_xxreal_0
s: 295
d: [2,809,1,2,1,1,40,268,615]
a: [3,824,1,2,1,809]
s: 824
i: [1,296,0,6,2,0,0,255] t1_subset
s: 296
i: [1,305,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 305
d: [2,810,1,2,1,1,40,268,625]
a: [3,829,1,2,1,810]
s: 829
i: [1,252,0,7,3,0,0,4] cc5_ordinal1
s: 252
d: [2,831,1,2,1,1,40,269,633]
a: [3,838,1,2,1,831]
s: 838
i: [1,266,0,7,2,0,0,2] involutiveness_k4_xcmplx_0
s: 266
i: [1,274,0,7,3,0,0,1] fc16_xreal_0
s: 274
d: [2,832,1,2,1,1,40,269,495]
a: [3,843,1,2,1,832]
s: 843
i: [1,275,0,7,3,0,0,1] fc16_xreal_0
s: 275
d: [2,833,1,2,1,1,40,269,615]
a: [3,848,1,2,1,833]
s: 848
i: [1,276,0,7,3,0,0,1] fc15_xreal_0
s: 276
i: [1,277,0,7,3,0,0,1] fc15_xreal_0
s: 277
d: [2,834,1,2,1,1,40,269,625]
a: [3,853,1,2,1,834]
s: 853
i: [1,278,0,7,3,0,0,1] fc6_xcmplx_0
s: 278
d: [2,877,1,2,1,1,40,271,645]
a: [3,883,1,2,1,877]
s: 883
i: [1,279,0,7,3,0,0,1] fc6_xcmplx_0
s: 279
i: [1,298,0,7,2,0,0,255] t3_subset
s: 298
d: [2,890,1,2,1,1,40,272,438]
a: [3,895,1,2,1,890]
s: 895
i: [1,299,0,7,3,0,0,2] t8_boole
s: 299
d: [2,902,1,2,1,1,40,273,438]
a: [3,907,1,2,1,902]
s: 907
i: [1,245,0,8,3,0,0,255] cc10_ordinal1
s: 245
i: [1,281,0,8,4,0,0,2] cc4_xxreal_0
s: 281
d: [2,939,2,5,2,3,40,929,930]
a: [3,951,2,2,1,939]
s: 951
r: 951
i: [1,283,0,8,4,0,0,2] cc6_xxreal_0
s: 283
a: [3,952,2,2,1,939]
s: 952
i: [1,285,0,8,4,0,0,2] cc8_xxreal_0
s: 285
i: [1,297,0,8,3,0,0,255] t2_subset
s: 297
d: [2,1181,2,6,3,1,40,281,774]
a: [3,1195,2,2,1,1181]
s: 1195
i: [1,302,0,9,3,0,0,255] t5_subset
s: 302
a: [3,1198,2,2,1,1181]
s: 1198
i: [1,301,0,10,3,0,0,255] t4_subset
s: 301
i: [1,311,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 311
a: [3,943,2,3,1,938]
s: 943
i: [1,259,0,11,5,0,0,1] t2_real
s: 259
d: [2,954,1,4,2,1,40,287,545]
d: [2,956,1,2,1,2,21,954,550]
a: [3,960,1,2,1,956]
s: 960
i: [1,260,0,11,5,0,0,1] t1_real
s: 260
i: [1,261,0,11,5,0,0,1] t8_real
s: 261
d: [2,1351,1,10,4,1,40,261,347]
a: [3,1361,1,3,1,1351]
s: 1361
d: [2,1375,2,0,0,2,52,358,1361,256]
e: 1375
d: [2,1376,2,2,1,1,40,1361,256]
d: [2,1378,2,0,0,2,21,1376,358]
e: 1378
r: 1361
i: [1,262,0,11,5,0,0,1] t7_real
s: 262
a: [3,1360,1,3,1,1351]
s: 1360
d: [2,1380,3,0,0,2,52,943,1360,254]
e: 1380
d: [2,1381,2,3,1,1,40,1360,254]
d: [2,1384,2,0,0,2,21,1381,943]
e: 1384
r: 943
i: [1,263,0,11,5,0,0,1] t3_real
s: 263
i: [1,264,0,11,5,0,0,1] t4_real
s: 264
a: [3,1365,1,3,1,1351]
s: 1365
d: [2,1392,3,0,0,3,52,1198,358,1365,275]
e: 1392
d: [2,1393,2,4,2,1,40,1365,275]
d: [2,1397,2,2,1,2,21,1393,358]
d: [2,1398,2,0,0,3,21,1397,1198]
e: 1398
r: 1195
r: 1198
r: 1365
i: [1,257,0,13,6,0,0,1] t5_real
s: 257
a: [3,1369,1,2,1,1351]
s: 1369
r: 929
r: 347
r: 930
r: 352
r: 1369
a: [3,351,0,3,1,224]
s: 351
a: [3,348,0,4,1,224]
s: 348
a: [3,1368,1,2,1,1351]
s: 1368
i: [1,258,0,13,6,0,0,1] t6_real
s: 258
d: [2,1342,1,8,4,1,40,260,352]
d: [2,1400,1,6,3,2,21,1342,358]
a: [3,1403,1,2,1,1400]
s: 1403
r: 947
r: 1403
i: [1,309,0,14,4,0,0,2] existence_m2_subset_1
s: 309
i: [1,306,0,15,5,0,0,2] dt_m2_subset_1
s: 306
a: [3,1404,1,2,1,1400]
s: 1404
i: [1,307,0,15,5,0,0,2] redefinition_m2_subset_1
s: 307
a: [3,1194,2,2,1,1181]
s: 1194
i: [1,308,0,15,5,0,0,2] redefinition_m2_subset_1
s: 308
a: [3,1364,1,3,1,1351]
s: 1364
a: [3,1199,2,2,1,1181]
s: 1199
d: [2,1456,1,7,3,1,40,351,311]
d: [2,1460,1,5,2,2,21,1456,774]
a: [3,1464,1,3,1,1460]
s: 1464
a: [3,942,2,3,1,938]
s: 942
d: [2,1311,1,9,3,1,40,311,347]
a: [3,1319,1,4,1,1311]
s: 1319
d: [2,1457,1,5,2,2,44,351,875]
s: 1457
a: [3,946,2,3,1,938]
s: 946
d: [2,1545,3,0,0,3,52,952,755,946,278]
e: 1545
d: [2,1546,3,4,2,1,40,946,278]
d: [2,1551,3,2,1,2,21,1546,755]
d: [2,1552,3,0,0,3,21,1551,952]
e: 1552
r: 952
r: 946
r: 1319
a: [3,1318,1,4,1,1311]
s: 1318
s: 952
d: [2,1469,1,6,2,2,44,348,875]
s: 1469
d: [2,955,1,4,2,1,40,287,580]
d: [2,962,1,2,1,2,21,955,585]
a: [3,966,1,2,1,962]
s: 966
d: [2,1250,3,2,1,1,40,951,269]
a: [3,1256,3,2,1,1250]
s: 1256
r: 1256
d: [2,968,1,4,2,1,40,289,453]
a: [3,973,1,2,1,968]
s: 973
a: [3,1255,3,2,1,1250]
s: 1255
a: [3,978,1,2,1,968]
s: 978
d: [2,969,1,4,2,1,40,289,565]
d: [2,980,1,2,1,2,21,969,570]
a: [3,984,1,2,1,980]
s: 984
d: [2,1251,3,2,1,1,40,951,268]
a: [3,1261,3,2,1,1251]
s: 1261
d: [2,1608,4,0,0,2,52,1255,1261,246]
e: 1608
d: [2,1612,4,2,1,1,40,1261,246]
d: [2,1614,4,0,0,2,21,1612,1255]
e: 1614
r: 1261
d: [2,970,1,4,2,1,40,289,590]
d: [2,986,1,2,1,2,21,970,595]
a: [3,990,1,2,1,986]
s: 990
d: [2,1252,3,2,1,1,40,951,267]
a: [3,1266,3,2,1,1252]
s: 1266
d: [2,993,1,4,2,1,40,290,565]
d: [2,995,1,2,1,2,21,993,570]
a: [3,999,1,2,1,995]
s: 999
d: [2,994,1,4,2,1,40,290,590]
d: [2,1001,1,2,1,2,21,994,595]
a: [3,1005,1,2,1,1001]
s: 1005
s: 947
d: [2,1070,1,4,2,1,40,252,438]
a: [3,1082,1,2,1,1070]
s: 1082
d: [2,1382,2,3,1,1,40,1360,253]
a: [3,1390,2,3,1,1382]
s: 1390
d: [2,1170,1,4,2,1,40,245,859]
a: [3,1175,1,2,1,1170]
s: 1175
d: [2,1184,1,6,3,1,40,281,490]
a: [3,1230,1,2,1,1184]
s: 1230
r: 1230
d: [2,1395,2,6,2,1,40,1365,289]
a: [3,1418,2,3,1,1395]
s: 1418
d: [2,1646,3,0,0,3,52,1194,358,1418,277]
e: 1646
d: [2,1647,3,4,2,1,40,1418,277]
d: [2,1650,3,2,1,2,21,1647,358]
d: [2,1651,3,0,0,3,21,1650,1194]
e: 1651
i: [1,221,0,7,2,1,0,0] t58_xreal_1
d: [2,876,1,2,1,1,40,271,639]
d: [2,1007,2,2,1,1,40,814,249]
d: [2,1009,2,2,1,1,40,814,247]
d: [2,1302,3,4,2,1,40,1195,289]
d: [2,1303,3,2,1,2,21,1302,774]
d: [2,1328,1,8,4,1,40,259,352]
d: [2,1412,1,6,3,2,21,1328,358]
d: [2,1451,2,4,2,1,40,1369,287]
d: [2,1453,2,2,1,2,21,1451,633]
d: [2,1455,1,8,4,1,40,351,261]
d: [2,1458,1,6,3,2,21,1455,358]
d: [2,1491,2,4,2,1,40,1403,290]
d: [2,1494,2,2,1,2,21,1491,633]
d: [2,1480,1,13,5,1,40,258,348]
d: [2,1496,1,10,4,2,21,1480,1360]
d: [2,1541,2,4,2,3,40,1457,460]
d: [2,1555,2,2,1,4,21,1541,774]
f: 1651,1555,1551,1496,1494,1458,1453,1412,1397,1378,1351,1303,1009,1007,876,764,745,635,629,627,317,310,230,229,227,220,221,224
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t58_xreal_1
% SZS output start Proof for xreal_1__t58_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k4_xcmplx_0(X0),k6_numbers)) & ~(r1_xxreal_0(k4_xcmplx_0(X0),k6_numbers) & ~r1_xxreal_0(k6_numbers,X0)))) [input t58_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k4_xcmplx_0(X0),k6_numbers)) & ~(r1_xxreal_0(k4_xcmplx_0(X0),k6_numbers) & ~r1_xxreal_0(k6_numbers,X0)))) [negated conjecture 1]
12. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
13. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
16. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
18. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
20. ! [X0] : ((v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => (~v2_xxreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0)))) [input fc16_xreal_0]
27. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
29. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
35. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0))))) [input t6_real]
37. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
50. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
56. ! [X0] : ((v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => (~v3_xxreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0)))) [input fc15_xreal_0]
58. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
59. ! [X0] : ((v1_xcmplx_0(X0) & ~v1_xboole_0(X0)) => (v1_xcmplx_0(k4_xcmplx_0(X0)) & ~v1_xboole_0(k4_xcmplx_0(X0)))) [input fc6_xcmplx_0]
63. ? [X0] : v1_xxreal_0(X0) [input rc1_xxreal_0]
73. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
77. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => r1_xxreal_0(X0,X0)) [input reflexivity_r1_xxreal_0]
79. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
82. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
83. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
84. ! [X0] : (v1_xreal_0(X0) => (v1_xreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0)))) [input fc3_xreal_0]
105. ? [X0] : (((r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k4_xcmplx_0(X0),k6_numbers)) | (r1_xxreal_0(k4_xcmplx_0(X0),k6_numbers) & ~r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X0)) [ennf transformation 2]
108. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
110. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 13]
113. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 82]
114. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 83]
115. ! [X0] : ((v1_xreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0))) | ~v1_xreal_0(X0)) [ennf transformation 84]
118. ! [X0] : (! [X1] : (((v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0)) | ~r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 35]
119. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 118]
120. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 29]
121. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 120]
124. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 37]
125. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 124]
137. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 73]
138. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 18]
141. ! [X0] : ((~v2_xxreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0))) | (~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 20]
142. ! [X0] : ((~v2_xxreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0))) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 141]
143. ! [X0] : ((~v3_xxreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0))) | (~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 56]
144. ! [X0] : ((~v3_xxreal_0(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0))) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 143]
145. ! [X0] : ((v1_xcmplx_0(k4_xcmplx_0(X0)) & ~v1_xboole_0(k4_xcmplx_0(X0))) | (~v1_xcmplx_0(X0) | v1_xboole_0(X0))) [ennf transformation 59]
146. ! [X0] : ((v1_xcmplx_0(k4_xcmplx_0(X0)) & ~v1_xboole_0(k4_xcmplx_0(X0))) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0)) [flattening 145]
155. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 16]
156. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 155]
157. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 50]
158. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 157]
177. ! [X0,X1] : (r1_xxreal_0(X0,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 77]
178. ! [X0,X1] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 177]
181. ? [X0] : (((r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k4_xcmplx_0(X0),k6_numbers)) | (r1_xxreal_0(k4_xcmplx_0(X0),k6_numbers) & ~r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X0)) => (((r1_xxreal_0(k6_numbers,sK0) & ~r1_xxreal_0(k4_xcmplx_0(sK0),k6_numbers)) | (r1_xxreal_0(k4_xcmplx_0(sK0),k6_numbers) & ~r1_xxreal_0(k6_numbers,sK0))) & v1_xreal_0(sK0)) [choice axiom]
182. ((r1_xxreal_0(k6_numbers,sK0) & ~r1_xxreal_0(k4_xcmplx_0(sK0),k6_numbers)) | (r1_xxreal_0(k4_xcmplx_0(sK0),k6_numbers) & ~r1_xxreal_0(k6_numbers,sK0))) & v1_xreal_0(sK0) [skolemisation 105,181]
198. ? [X0] : v1_xxreal_0(X0) => v1_xxreal_0(sK8) [choice axiom]
199. v1_xxreal_0(sK8) [skolemisation 63,198]
220. v1_xreal_0(sK0) [cnf transformation 182] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
221. ~r1_xxreal_0(k4_xcmplx_0(sK0),k6_numbers) | ~r1_xxreal_0(k6_numbers,sK0) [cnf transformation 182] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
224. r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(k4_xcmplx_0(sK0),k6_numbers) [cnf transformation 182] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
227. v1_xboole_0(k1_xboole_0) [cnf transformation 58] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
229. k6_numbers = k1_xboole_0 [cnf transformation 79] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
230. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 27] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
247. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 108] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
249. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 110] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
253. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 113] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
254. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 114] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
256. v1_xreal_0(k4_xcmplx_0(X0)) | ~v1_xreal_0(X0) [cnf transformation 115] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
258. ~r1_xxreal_0(X0,X1) | v2_xxreal_0(X1) | v1_xboole_0(X0) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 119] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
259. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 121] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
261. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 125] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
270. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 137] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
271. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 138] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
275. ~v2_xxreal_0(k4_xcmplx_0(X0)) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 142] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
277. ~v3_xxreal_0(k4_xcmplx_0(X0)) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 144] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
278. ~v1_xboole_0(k4_xcmplx_0(X0)) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) [cnf transformation 146] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
287. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 156] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
289. ~v2_xxreal_0(X0) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 156] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
290. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 158] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
310. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 178] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
317. v1_xxreal_0(sK8) [cnf transformation 199] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
346. 1 <=> r1_xxreal_0(k4_xcmplx_0(sK0),k6_numbers) [avatar definition]
347. ~r1_xxreal_0(k4_xcmplx_0(sK0),k6_numbers) <- (~1) [avatar component clause 346] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
348. r1_xxreal_0(k4_xcmplx_0(sK0),k6_numbers) <- (1) [avatar component clause 346] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
350. 2 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
351. ~r1_xxreal_0(k6_numbers,sK0) <- (~2) [avatar component clause 350] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
352. r1_xxreal_0(k6_numbers,sK0) <- (2) [avatar component clause 350] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
353. 1 | 2 [avatar split clause 224,350,346]
354. ~2 | ~1 [avatar split clause 221,346,350]
356. 3 <=> v1_xreal_0(sK0) [avatar definition]
358. v1_xreal_0(sK0) <- (3) [avatar component clause 356] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
359. 3 [avatar split clause 220,356]
371. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
373. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 371] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
374. 6 [avatar split clause 227,371]
381. 8 <=> k6_numbers = k1_xboole_0 [avatar definition]
383. k6_numbers = k1_xboole_0 <- (8) [avatar component clause 381] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
384. 8 [avatar split clause 229,381]
386. 9 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
388. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (9) [avatar component clause 386] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
389. 9 [avatar split clause 230,386]
456. 23 <=> ! [X1] : ~v1_xxreal_0(X1) [avatar definition]
457. ~v1_xxreal_0(X1) <- (23) [avatar component clause 456] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
459. 24 <=> ! [X0] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0)) [avatar definition]
460. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) <- (24) [avatar component clause 459] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
461. 23 | 24 [avatar split clause 310,459,456]
488. 30 <=> v1_xxreal_0(sK8) [avatar definition]
490. v1_xxreal_0(sK8) <- (30) [avatar component clause 488] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
491. 30 [avatar split clause 317,488]
627. $false <- (23, 30) [subsumption resolution 490,457] {a:0,w:0,sine:255,thAx:0,allAx:2,thDist:-2}
628. ~23 | ~30 [avatar contradiction clause 627]
629. v1_xboole_0(k6_numbers) <- (6, 8) [backward demodulation 373,383] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
631. 58 <=> v1_xboole_0(k6_numbers) [avatar definition]
633. v1_xboole_0(k6_numbers) <- (58) [avatar component clause 631] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
634. 58 | ~6 | ~8 [avatar split clause 629,381,371,631]
635. m1_subset_1(k6_numbers,k4_ordinal1) <- (8, 9) [forward demodulation 388,383] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
637. 59 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
639. m1_subset_1(k6_numbers,k4_ordinal1) <- (59) [avatar component clause 637] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
640. 59 | ~8 | ~9 [avatar split clause 635,386,381,637]
745. v1_xcmplx_0(sK0) <- (3) [resolution 253,358] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
753. 77 <=> v1_xcmplx_0(sK0) [avatar definition]
755. v1_xcmplx_0(sK0) <- (77) [avatar component clause 753] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
756. 77 | ~3 [avatar split clause 745,356,753]
764. v1_xxreal_0(sK0) <- (3) [resolution 254,358] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
772. 79 <=> v1_xxreal_0(sK0) [avatar definition]
774. v1_xxreal_0(sK0) <- (79) [avatar component clause 772] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
775. 79 | ~3 [avatar split clause 764,356,772]
812. 85 <=> v7_ordinal1(k6_numbers) [avatar definition]
814. v7_ordinal1(k6_numbers) <- (85) [avatar component clause 812] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
875. k6_numbers = X0 | ~v1_xboole_0(X0) <- (8) [forward demodulation 270,383] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
876. v7_ordinal1(k6_numbers) <- (59) [resolution 271,639] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
879. 85 | ~59 [avatar split clause 876,637,812]
945. 105 <=> v1_xboole_0(k4_xcmplx_0(sK0)) [avatar definition]
946. v1_xboole_0(k4_xcmplx_0(sK0)) <- (105) [avatar component clause 945] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
950. 106 <=> v1_xboole_0(sK0) [avatar definition]
1007. v1_xxreal_0(k6_numbers) <- (85) [resolution 814,249] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1009. v1_xreal_0(k6_numbers) <- (85) [resolution 814,247] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1013. 115 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1016. 115 | ~85 [avatar split clause 1007,812,1013]
1023. 117 <=> v1_xreal_0(k6_numbers) [avatar definition]
1026. 117 | ~85 [avatar split clause 1009,812,1023]
1193. 137 <=> v2_xxreal_0(sK0) [avatar definition]
1194. ~v2_xxreal_0(sK0) <- (~137) [avatar component clause 1193] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1195. v2_xxreal_0(sK0) <- (137) [avatar component clause 1193] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1197. 138 <=> v3_xxreal_0(sK0) [avatar definition]
1302. ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (137) [resolution 1195,289] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1303. ~v3_xxreal_0(sK0) <- (79, 137) [subsumption resolution 1302,774] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1304. ~138 | ~79 | ~137 [avatar split clause 1303,1193,772,1197]
1328. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (2) [resolution 259,352] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1351. v3_xxreal_0(k6_numbers) | v2_xxreal_0(k4_xcmplx_0(sK0)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k4_xcmplx_0(sK0)) <- (~1) [resolution 261,347] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1359. 160 <=> v1_xreal_0(k4_xcmplx_0(sK0)) [avatar definition]
1360. v1_xreal_0(k4_xcmplx_0(sK0)) <- (160) [avatar component clause 1359] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1361. ~v1_xreal_0(k4_xcmplx_0(sK0)) <- (~160) [avatar component clause 1359] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1363. 161 <=> v2_xxreal_0(k4_xcmplx_0(sK0)) [avatar definition]
1365. v2_xxreal_0(k4_xcmplx_0(sK0)) <- (161) [avatar component clause 1363] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1367. 162 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1369. v3_xxreal_0(k6_numbers) <- (162) [avatar component clause 1367] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1370. ~160 | ~117 | 161 | 162 | 1 [avatar split clause 1351,346,1367,1363,1023,1359]
1376. ~v1_xreal_0(sK0) <- (~160) [resolution 1361,256] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1378. $false <- (3, ~160) [subsumption resolution 1376,358] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1379. ~3 | 160 [avatar contradiction clause 1378]
1393. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (161) [resolution 1365,275] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1397. v3_xxreal_0(sK0) <- (3, 161) [subsumption resolution 1393,358] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1402. 165 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1403. v2_xxreal_0(k6_numbers) <- (165) [avatar component clause 1402] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1411. 138 | ~3 | ~161 [avatar split clause 1397,1363,356,1197]
1412. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 3) [subsumption resolution 1328,358] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1413. ~117 | 162 | ~138 | ~2 | ~3 [avatar split clause 1412,356,350,1197,1367,1023]
1417. 167 <=> v3_xxreal_0(k4_xcmplx_0(sK0)) [avatar definition]
1418. v3_xxreal_0(k4_xcmplx_0(sK0)) <- (167) [avatar component clause 1417] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1451. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (162) [resolution 1369,287] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1453. ~v1_xxreal_0(k6_numbers) <- (58, 162) [subsumption resolution 1451,633] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1454. ~115 | ~58 | ~162 [avatar split clause 1453,1367,631,1013]
1455. v3_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (~2) [resolution 351,261] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1457. ~r1_xxreal_0(X0,sK0) | ~v1_xboole_0(X0) <- (~2, 8) [superposition 351,875] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
1458. v3_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (~2, 3) [subsumption resolution 1455,358] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1459. ~117 | 165 | 138 | 2 | ~3 [avatar split clause 1458,356,350,1197,1402,1023]
1480. v2_xxreal_0(k6_numbers) | v1_xboole_0(k4_xcmplx_0(sK0)) | v3_xxreal_0(k4_xcmplx_0(sK0)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k4_xcmplx_0(sK0)) <- (1) [resolution 258,348] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1491. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (165) [resolution 1403,290] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1494. ~v1_xxreal_0(k6_numbers) <- (58, 165) [subsumption resolution 1491,633] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1495. ~115 | ~58 | ~165 [avatar split clause 1494,1402,631,1013]
1496. v2_xxreal_0(k6_numbers) | v1_xboole_0(k4_xcmplx_0(sK0)) | v3_xxreal_0(k4_xcmplx_0(sK0)) | ~v1_xreal_0(k6_numbers) <- (1, 160) [subsumption resolution 1480,1360] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1497. ~117 | 167 | 105 | 165 | ~1 | ~160 [avatar split clause 1496,1359,346,1402,945,1417,1023]
1541. ~v1_xboole_0(sK0) | ~v1_xxreal_0(sK0) <- (~2, 8, 24) [resolution 1457,460] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1546. ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (105) [resolution 946,278] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1551. v1_xboole_0(sK0) <- (77, 105) [subsumption resolution 1546,755] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1554. 106 | ~77 | ~105 [avatar split clause 1551,945,753,950]
1555. ~v1_xboole_0(sK0) <- (~2, 8, 24, 79) [subsumption resolution 1541,774] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1556. ~106 | 2 | ~8 | ~24 | ~79 [avatar split clause 1555,772,459,381,350,950]
1647. ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (167) [resolution 1418,277] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1650. v2_xxreal_0(sK0) <- (3, 167) [subsumption resolution 1647,358] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1651. $false <- (3, ~137, 167) [subsumption resolution 1650,1194] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1652. ~3 | 137 | ~167 [avatar contradiction clause 1651]
1653. $false [avatar sat refutation 353,354,359,374,384,389,461,491,628,634,640,756,775,879,1016,1026,1304,1370,1379,1411,1413,1454,1459,1495,1497,1554,1556,1652] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1188,thDist:-1188}
% SZS output end Proof for xreal_1__t58_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 92
% >>> Preprocessing
% Introduced skolems: 19
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 125
% Generated clauses: 681
% Active clauses: 184
% Passive clauses: 281
% Final active clauses: 163
% Final passive clauses: 70
% Main loop iterations started: 202
% 
% >>> Simplifying Inferences
% Duplicate literals: 27
% Fw subsumption resolutions: 65
% Fw demodulations: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 20
% Equational tautologies: 1
% Forward subsumptions: 80
% 
% >>> Generating Inferences
% Binary resolution: 231
% Unit resulting resolution: 6
% Factoring: 1
% Forward superposition: 17
% Backward superposition: 12
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 47
% Split components: 118
% Sat splitting refutations: 13
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 224
% SAT solver unit clauses: 56
% SAT solver binary clauses: 93
% 
% Memory used [KB]: 6012
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% non unit clause index maintenance: 0.001 s
% binary resolution index maintenance: 0.001 s
% unit resulting resolution: 0.002 s
% other: 0.009 s ( own 0.003 s ) 

