i: [1,183,0,2,1,1,0,0] t6_xxreal_1
a: [3,296,0,2,1,183]
s: 296
i: [1,182,0,2,1,1,0,0] t6_xxreal_1
a: [3,301,0,2,1,182]
s: 301
i: [1,186,0,3,1,1,0,0] t6_xxreal_1
a: [3,281,0,3,1,186]
s: 281
i: [1,185,0,5,1,1,0,0] t6_xxreal_1
a: [3,286,0,5,1,185]
s: 286
i: [1,184,0,5,1,1,0,0] t6_xxreal_1
a: [3,291,0,5,1,184]
s: 291
i: [1,187,0,2,1,0,0,3] fc5_numbers
a: [3,306,0,2,1,187]
s: 306
i: [1,188,0,2,1,0,0,2] fc2_membered
a: [3,311,0,2,1,188]
s: 311
i: [1,189,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,316,0,2,1,189]
s: 316
i: [1,190,0,2,1,0,0,255] fc1_xboole_0
a: [3,321,0,2,1,190]
s: 321
i: [1,251,0,2,1,0,0,2] rc2_xboole_0
a: [3,338,0,2,1,251]
s: 338
i: [1,252,0,2,1,0,0,255] rc1_xreal_0
a: [3,343,0,2,1,252]
s: 343
i: [1,253,0,2,1,0,0,2] rc1_xboole_0
a: [3,348,0,2,1,253]
s: 348
i: [1,254,0,2,1,0,0,1] rc1_xxreal_0
a: [3,353,0,2,1,254]
s: 353
i: [1,256,0,2,1,0,0,255] rc1_membered
a: [3,358,0,2,1,256]
s: 358
i: [1,255,0,2,1,0,0,255] rc1_membered
a: [3,363,0,2,1,255]
s: 363
i: [1,258,0,2,1,0,0,255] rc2_membered
a: [3,368,0,2,1,258]
s: 368
i: [1,257,0,2,1,0,0,255] rc2_membered
a: [3,373,0,2,1,257]
s: 373
i: [1,260,0,2,1,0,0,255] rc3_membered
a: [3,378,0,2,1,260]
s: 378
i: [1,259,0,2,1,0,0,255] rc3_membered
a: [3,383,0,2,1,259]
s: 383
i: [1,263,0,2,1,0,0,255] rc3_xreal_0
a: [3,388,0,2,1,263]
s: 388
i: [1,262,0,2,1,0,0,255] rc3_xreal_0
a: [3,393,0,2,1,262]
s: 393
i: [1,261,0,2,1,0,0,255] rc3_xreal_0
a: [3,398,0,2,1,261]
s: 398
i: [1,266,0,2,1,0,0,255] rc2_xreal_0
a: [3,403,0,2,1,266]
s: 403
i: [1,265,0,2,1,0,0,255] rc2_xreal_0
a: [3,408,0,2,1,265]
s: 408
i: [1,264,0,2,1,0,0,255] rc2_xreal_0
a: [3,413,0,2,1,264]
s: 413
i: [1,269,0,2,1,0,0,255] rc4_xreal_0
a: [3,418,0,2,1,269]
s: 418
i: [1,268,0,2,1,0,0,255] rc4_xreal_0
a: [3,423,0,2,1,268]
s: 423
i: [1,267,0,2,1,0,0,255] rc4_xreal_0
a: [3,428,0,2,1,267]
s: 428
i: [1,271,0,2,1,0,0,2] rc4_xxreal_0
a: [3,433,0,2,1,271]
s: 433
i: [1,270,0,2,1,0,0,2] rc4_xxreal_0
a: [3,438,0,2,1,270]
s: 438
i: [1,273,0,2,1,0,0,255] rc3_xxreal_0
a: [3,443,0,2,1,273]
s: 443
i: [1,272,0,2,1,0,0,255] rc3_xxreal_0
a: [3,448,0,2,1,272]
s: 448
i: [1,275,0,2,1,0,0,255] rc2_xxreal_0
a: [3,453,0,2,1,275]
s: 453
i: [1,274,0,2,1,0,0,255] rc2_xxreal_0
a: [3,458,0,2,1,274]
s: 458
i: [1,191,0,3,1,0,0,255] d2_xboole_0
a: [3,326,0,3,1,191]
s: 326
r: 316
i: [1,192,0,4,2,0,0,255] cc2_membered
s: 192
i: [1,193,0,4,2,0,0,255] cc3_membered
s: 193
i: [1,194,0,4,2,0,0,255] cc2_xreal_0
s: 194
i: [1,195,0,4,2,0,0,255] cc2_xxreal_0
s: 195
i: [1,196,0,4,2,0,0,255] cc4_membered
s: 196
i: [1,199,0,4,2,0,0,255] cc1_membered
s: 199
d: [2,461,1,2,1,1,40,199,358]
a: [3,467,1,2,1,461]
s: 467
i: [1,201,0,4,2,0,0,255] cc4_xreal_0
s: 201
d: [2,462,1,2,1,1,40,199,368]
a: [3,472,1,2,1,462]
s: 472
i: [1,202,0,4,2,0,0,255] cc18_membered
s: 202
i: [1,235,0,4,1,0,0,2] existence_m1_subset_1
s: 235
d: [2,463,1,2,1,1,40,199,378]
a: [3,477,1,2,1,463]
s: 477
i: [1,203,0,5,2,0,0,255] t6_boole
s: 203
d: [2,485,1,2,1,1,40,201,343]
a: [3,492,1,2,1,485]
s: 492
i: [1,213,0,5,2,0,0,3] cc1_xxreal_0
s: 213
i: [1,232,0,5,2,0,0,1] t7_boole
s: 232
d: [2,538,1,4,1,1,40,232,291]
a: [3,542,1,4,1,538]
s: 542
i: [1,237,0,7,3,0,0,2] reflexivity_r1_xxreal_0
a: [3,333,0,5,2,237]
s: 333
d: [2,500,1,2,1,1,40,202,321]
a: [3,507,1,2,1,500]
s: 507
i: [1,219,0,6,3,0,0,255] cc7_xxreal_0
s: 219
i: [1,220,0,6,3,0,0,255] cc7_xxreal_0
s: 220
d: [2,501,1,2,1,1,40,202,348]
a: [3,512,1,2,1,501]
s: 512
i: [1,223,0,6,3,0,0,255] cc5_xxreal_0
s: 223
d: [2,502,1,2,1,1,40,202,428]
a: [3,517,1,2,1,502]
s: 517
i: [1,225,0,6,3,0,0,255] cc5_xxreal_0
s: 225
i: [1,226,0,6,3,0,0,255] cc3_xxreal_0
s: 226
d: [2,503,1,2,1,1,40,202,438]
a: [3,522,1,2,1,503]
s: 522
i: [1,228,0,6,3,0,0,255] cc3_xxreal_0
s: 228
d: [2,551,1,4,2,1,40,219,408]
d: [2,553,1,2,1,2,21,551,413]
a: [3,557,1,2,1,553]
s: 557
i: [1,230,0,6,2,0,0,1] t1_subset
s: 230
i: [1,236,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 236
d: [2,611,1,5,1,1,40,230,291]
a: [3,615,1,5,1,611]
s: 615
i: [1,197,0,7,3,0,0,255] cc14_membered
s: 197
d: [2,624,2,6,2,1,40,197,615]
a: [3,628,2,4,1,624]
s: 628
i: [1,198,0,7,3,0,0,2] cc13_membered
s: 198
i: [1,200,0,7,3,0,0,255] cc17_membered
s: 200
d: [2,635,2,6,2,1,40,198,615]
a: [3,639,2,4,1,635]
s: 639
i: [1,231,0,7,3,0,0,1] t8_boole
s: 231
d: [2,646,2,6,2,1,40,200,615]
a: [3,650,2,4,1,646]
s: 650
i: [1,233,0,7,2,0,0,1] t2_tarski
s: 233
i: [1,215,0,8,4,0,0,255] cc6_xxreal_0
s: 215
d: [2,659,1,6,3,1,40,215,301]
a: [3,672,1,2,1,659]
s: 672
r: 672
i: [1,217,0,8,4,0,0,255] cc4_xxreal_0
s: 217
a: [3,671,1,2,1,659]
s: 671
i: [1,221,0,8,4,0,0,255] cc8_xxreal_0
s: 221
i: [1,229,0,8,3,0,0,1] t2_subset
s: 229
a: [3,676,1,2,1,659]
s: 676
i: [1,238,0,8,3,0,0,1] fc1_xxreal_1
s: 238
d: [2,733,3,0,0,3,52,301,296,639,238]
e: 733
d: [2,734,3,4,2,1,40,238,639]
d: [2,736,3,2,1,2,21,734,296]
d: [2,737,3,0,0,3,21,736,301]
e: 737
r: 639
a: [3,679,1,2,1,659]
s: 679
i: [1,239,0,8,3,0,0,1] fc3_xxreal_1
s: 239
i: [1,234,0,10,3,0,0,1] t2_tarski
s: 234
d: [2,660,1,6,3,1,40,215,296]
a: [3,685,1,2,1,660]
s: 685
r: 685
i: [1,240,0,10,4,0,0,2] connectedness_r1_xxreal_0
s: 240
a: [3,684,1,2,1,660]
s: 684
i: [1,204,0,11,3,0,0,255] d3_xxreal_1
s: 204
i: [1,205,0,11,3,0,0,255] d1_xxreal_1
s: 205
a: [3,689,1,2,1,660]
s: 689
i: [1,206,0,13,5,0,0,2] t1_xxreal_0
s: 206
a: [3,692,1,2,1,660]
s: 692
i: [1,207,0,14,5,0,0,1] t2_xxreal_1
s: 207
i: [1,208,0,14,5,0,0,1] t2_xxreal_1
s: 208
a: [3,643,2,2,1,635]
s: 643
i: [1,210,0,14,5,0,0,1] t1_xxreal_1
s: 210
d: [2,758,3,6,3,1,40,643,215]
a: [3,762,3,2,1,758]
s: 762
r: 762
i: [1,211,0,14,5,0,0,1] t1_xxreal_1
s: 211
i: [1,241,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 241
a: [3,761,3,2,1,758]
s: 761
i: [1,242,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 242
a: [3,766,3,2,1,758]
s: 766
i: [1,243,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 243
i: [1,244,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 244
a: [3,769,3,2,1,758]
s: 769
i: [1,246,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 246
d: [2,772,1,9,4,1,40,210,291]
d: [2,774,1,7,3,2,21,772,296]
d: [2,775,1,5,2,3,21,774,301]
d: [2,776,1,3,1,4,21,775,643]
a: [3,780,1,3,1,776]
s: 780
i: [1,247,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 247
i: [1,248,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 248
d: [2,788,1,9,4,1,40,211,291]
d: [2,790,1,7,3,2,21,788,296]
d: [2,791,1,5,2,3,21,790,301]
d: [2,792,1,3,1,4,21,791,643]
a: [3,796,1,3,1,792]
s: 796
i: [1,249,0,15,4,0,0,255] fraenkel_a_2_0_xxreal_1
s: 249
d: [2,852,2,10,4,1,40,780,206]
d: [2,853,2,7,3,2,21,852,281]
d: [2,854,2,5,2,3,21,853,301]
d: [2,855,2,3,1,4,21,854,643]
a: [3,859,2,3,1,855]
s: 859
i: [1,209,0,17,6,0,0,1] t2_xxreal_1
s: 209
d: [2,884,3,0,0,6,52,643,301,296,859,796,286,209]
e: 884
d: [2,885,1,12,5,1,40,209,286]
d: [2,894,1,9,4,2,21,885,796]
d: [2,895,1,6,3,3,21,894,859]
d: [2,896,1,4,2,4,21,895,296]
d: [2,897,1,2,1,5,21,896,301]
d: [2,898,1,0,0,6,21,897,643]
e: 898
f: 898,855,792,776,737,635,611,182,183,184,185,186
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_1__t6_xxreal_1
% SZS output start Proof for xxreal_1__t6_xxreal_1
1. ! [X0] : ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ~(X0 != X1 & ~r2_hidden(X0,k3_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2))))) [input t6_xxreal_1]
2. ~! [X0] : ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ~(X0 != X1 & ~r2_hidden(X0,k3_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2))))) [negated conjecture 1]
5. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xxreal_0(X1))) [input cc13_membered]
40. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v2_membered(k1_xxreal_1(X0,X1))) [input fc1_xxreal_1]
60. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
61. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input t1_xxreal_0]
62. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k1_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0)))))) [input t1_xxreal_1]
65. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k3_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1)))))) [input t2_xxreal_1]
69. ~! [X0,X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ~(X0 != X1 & ~r2_hidden(X0,k3_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2))))) [flattening 2]
71. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 60]
85. ? [X0,X1] : (? [X2] : ((X0 != X1 & ~r2_hidden(X0,k3_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2))) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) [ennf transformation 69]
86. ? [X0,X1] : (? [X2] : (X0 != X1 & ~r2_hidden(X0,k3_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) [flattening 85]
93. ! [X0] : (! [X1] : (v1_xxreal_0(X1) | ~m1_subset_1(X1,X0)) | ~v2_membered(X0)) [ennf transformation 5]
101. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 61]
102. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 101]
103. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X0,k3_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 65]
104. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X0,k1_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 62]
120. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 71]
127. ! [X0,X1] : (v2_membered(k1_xxreal_1(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 40]
128. ! [X0,X1] : (v2_membered(k1_xxreal_1(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 127]
137. ? [X0,X1] : (? [X2] : (X0 != X1 & ~r2_hidden(X0,k3_xxreal_1(X1,X2)) & r2_hidden(X0,k1_xxreal_1(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (sK0 != sK1 & ~r2_hidden(sK0,k3_xxreal_1(sK1,X2)) & r2_hidden(sK0,k1_xxreal_1(sK1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
138. ? [X2] : (sK0 != sK1 & ~r2_hidden(sK0,k3_xxreal_1(sK1,X2)) & r2_hidden(sK0,k1_xxreal_1(sK1,X2)) & v1_xxreal_0(X2)) => (sK0 != sK1 & ~r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) & r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) & v1_xxreal_0(sK2)) [choice axiom]
139. (sK0 != sK1 & ~r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) & r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1) [skolemisation 86,138,137]
140. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k3_xxreal_1(X1,X2)) | (~r1_xxreal_0(X0,X2) | r1_xxreal_0(X0,X1))) & ((r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1)) | ~r2_hidden(X0,k3_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 103]
141. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k3_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X2) | r1_xxreal_0(X0,X1)) & ((r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1)) | ~r2_hidden(X0,k3_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 140]
142. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k1_xxreal_1(X1,X2)) | (~r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X0))) & ((r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0)) | ~r2_hidden(X0,k1_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 104]
143. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k1_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X0)) & ((r1_xxreal_0(X0,X2) & r1_xxreal_0(X1,X0)) | ~r2_hidden(X0,k1_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 142]
182. v1_xxreal_0(sK1) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
183. v1_xxreal_0(sK2) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
184. r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) [cnf transformation 139] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
185. ~r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) [cnf transformation 139] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
186. sK0 != sK1 [cnf transformation 139] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
198. ~m1_subset_1(X1,X0) | v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 93] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
206. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 102] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
209. r2_hidden(X0,k3_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X2) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 141] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
210. ~r2_hidden(X0,k1_xxreal_1(X1,X2)) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 143] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
211. ~r2_hidden(X0,k1_xxreal_1(X1,X2)) | r1_xxreal_0(X0,X2) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 143] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
230. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 120] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
238. v2_membered(k1_xxreal_1(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 128] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
279. 1 <=> sK0 = sK1 [avatar definition]
281. sK0 != sK1 <- (~1) [avatar component clause 279] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
282. ~1 [avatar split clause 186,279]
284. 2 <=> r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) [avatar definition]
286. ~r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) <- (~2) [avatar component clause 284] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
287. ~2 [avatar split clause 185,284]
289. 3 <=> r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) [avatar definition]
291. r2_hidden(sK0,k1_xxreal_1(sK1,sK2)) <- (3) [avatar component clause 289] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
292. 3 [avatar split clause 184,289]
294. 4 <=> v1_xxreal_0(sK2) [avatar definition]
296. v1_xxreal_0(sK2) <- (4) [avatar component clause 294] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
297. 4 [avatar split clause 183,294]
299. 5 <=> v1_xxreal_0(sK1) [avatar definition]
301. v1_xxreal_0(sK1) <- (5) [avatar component clause 299] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
302. 5 [avatar split clause 182,299]
611. m1_subset_1(sK0,k1_xxreal_1(sK1,sK2)) <- (3) [resolution 230,291] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
613. 61 <=> m1_subset_1(sK0,k1_xxreal_1(sK1,sK2)) [avatar definition]
615. m1_subset_1(sK0,k1_xxreal_1(sK1,sK2)) <- (61) [avatar component clause 613] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
616. 61 | ~3 [avatar split clause 611,289,613]
635. v1_xxreal_0(sK0) | ~v2_membered(k1_xxreal_1(sK1,sK2)) <- (61) [resolution 198,615] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
637. 65 <=> v2_membered(k1_xxreal_1(sK1,sK2)) [avatar definition]
639. ~v2_membered(k1_xxreal_1(sK1,sK2)) <- (~65) [avatar component clause 637] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
641. 66 <=> v1_xxreal_0(sK0) [avatar definition]
643. v1_xxreal_0(sK0) <- (66) [avatar component clause 641] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
644. ~65 | 66 | ~61 [avatar split clause 635,613,641,637]
734. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~65) [resolution 238,639] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
736. ~v1_xxreal_0(sK1) <- (4, ~65) [subsumption resolution 734,296] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
737. $false <- (4, 5, ~65) [subsumption resolution 736,301] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
738. ~4 | ~5 | 65 [avatar contradiction clause 737]
772. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3) [resolution 210,291] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
774. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3, 4) [subsumption resolution 772,296] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
775. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) <- (3, 4, 5) [subsumption resolution 774,301] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
776. r1_xxreal_0(sK1,sK0) <- (3, 4, 5, 66) [subsumption resolution 775,643] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
778. 86 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
780. r1_xxreal_0(sK1,sK0) <- (86) [avatar component clause 778] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
781. 86 | ~3 | ~4 | ~5 | ~66 [avatar split clause 776,641,299,294,289,778]
788. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3) [resolution 211,291] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
790. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3, 4) [subsumption resolution 788,296] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
791. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK0) <- (3, 4, 5) [subsumption resolution 790,301] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
792. r1_xxreal_0(sK0,sK2) <- (3, 4, 5, 66) [subsumption resolution 791,643] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
794. 88 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
796. r1_xxreal_0(sK0,sK2) <- (88) [avatar component clause 794] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
797. 88 | ~3 | ~4 | ~5 | ~66 [avatar split clause 792,641,299,294,289,794]
852. sK0 = sK1 | ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (86) [resolution 780,206] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
853. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~1, 86) [subsumption resolution 852,281] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
854. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK0) <- (~1, 5, 86) [subsumption resolution 853,301] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
855. ~r1_xxreal_0(sK0,sK1) <- (~1, 5, 66, 86) [subsumption resolution 854,643] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
857. 95 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
859. ~r1_xxreal_0(sK0,sK1) <- (~95) [avatar component clause 857] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
860. ~95 | 1 | ~5 | ~66 | ~86 [avatar split clause 855,778,641,299,279,857]
885. ~r1_xxreal_0(sK0,sK2) | r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2) [resolution 209,286] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
894. r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, 88) [subsumption resolution 885,796] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
895. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, 88, ~95) [subsumption resolution 894,859] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
896. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, 4, 88, ~95) [subsumption resolution 895,296] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
897. ~v1_xxreal_0(sK0) <- (~2, 4, 5, 88, ~95) [subsumption resolution 896,301] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
898. $false <- (~2, 4, 5, 66, 88, ~95) [subsumption resolution 897,643] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
899. 2 | ~4 | ~5 | ~66 | ~88 | 95 [avatar contradiction clause 898]
900. $false [avatar sat refutation 282,287,292,297,302,616,644,738,781,797,860,899] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:464,thDist:-464}
% SZS output end Proof for xxreal_1__t6_xxreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 67
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 5
% 
% >>> Saturation
% Initial clauses: 94
% Generated clauses: 333
% Active clauses: 110
% Passive clauses: 167
% Final active clauses: 105
% Final passive clauses: 52
% Main loop iterations started: 115
% 
% >>> Simplifying Inferences
% Duplicate literals: 37
% Fw subsumption resolutions: 31
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 1
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 93
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 14
% Split components: 33
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 89
% SAT solver unit clauses: 36
% SAT solver binary clauses: 26
% 
% Memory used [KB]: 9978
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s ( own 0.0000 s ) 
% sine selection: 0.001 s
% unit clause index maintenance: 0.001 s
% backward subsumption index maintenance: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

