i: [1,52,0,2,1,1,0,0] t9_funct_4
a: [3,86,0,2,1,52]
s: 86
i: [1,51,0,2,1,1,0,0] t9_funct_4
a: [3,91,0,2,1,51]
s: 91
i: [1,50,0,2,1,1,0,0] t9_funct_4
a: [3,96,0,2,1,50]
s: 96
i: [1,49,0,2,1,1,0,0] t9_funct_4
a: [3,101,0,2,1,49]
s: 101
i: [1,53,0,3,1,1,0,0] t9_funct_4
a: [3,81,0,3,1,53]
s: 81
i: [1,72,0,2,1,0,0,1] rc1_funct_1
a: [3,106,0,2,1,72]
s: 106
i: [1,71,0,2,1,0,0,1] rc1_funct_1
a: [3,111,0,2,1,71]
s: 111
i: [1,54,0,11,1,1,0,0] t9_funct_4
a: [3,76,0,11,1,54]
s: 76
i: [1,62,0,3,1,0,0,1] reflexivity_r1_tarski
s: 62
i: [1,61,0,4,1,0,0,255] existence_m1_subset_1
s: 61
i: [1,63,0,6,2,0,0,1] dt_k5_relat_1
s: 63
i: [1,64,0,6,2,0,0,1] dt_k6_relat_1
s: 64
i: [1,59,0,7,2,0,0,255] t3_subset
s: 59
d: [2,113,1,5,1,0,40,59,61]
s: 113
i: [1,60,0,7,2,0,0,255] t3_subset
s: 60
i: [1,55,0,8,3,0,0,255] cc2_relat_1
s: 55
d: [2,115,1,6,2,0,40,55,61]
s: 115
i: [1,67,0,8,3,0,0,1] fc8_funct_1
s: 67
d: [2,116,1,7,3,0,40,55,60]
s: 116
i: [1,68,0,8,3,0,0,1] fc8_funct_1
s: 68
i: [1,69,0,8,3,0,0,1] fc9_funct_1
s: 69
i: [1,70,0,8,3,0,0,1] fc9_funct_1
s: 70
i: [1,56,0,10,4,0,0,255] cc3_funct_1
s: 56
d: [2,120,1,8,3,0,40,56,61]
s: 120
i: [1,65,0,11,2,0,0,255] rd7_relat_1
s: 65
i: [1,66,0,11,2,0,0,255] rd5_relat_1
s: 66
d: [2,121,1,9,4,0,40,56,60]
s: 121
i: [1,57,0,14,4,0,0,1] t76_relat_1
s: 57
d: [2,131,0,12,3,0,21,57,116]
s: 131
i: [1,58,0,14,4,0,0,1] t101_relat_1
s: 58
d: [2,149,0,12,3,0,21,58,116]
s: 149
d: [2,132,1,11,2,1,40,131,76]
a: [3,140,1,4,1,132]
s: 140
d: [2,157,2,0,0,2,52,91,140,64]
e: 157
d: [2,158,2,2,1,1,40,140,64]
d: [2,160,2,0,0,2,21,158,91]
e: 160
r: 140
a: [3,139,1,4,1,132]
s: 139
a: [3,144,1,7,1,132]
s: 144
d: [2,162,2,0,0,3,52,91,81,144,149]
e: 162
d: [2,163,2,5,2,1,40,144,149]
d: [2,165,2,2,1,2,21,163,81]
d: [2,166,2,0,0,3,21,165,91]
e: 166
f: 166,160,132,51,53,54
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t9_funct_4
% SZS output start Proof for funct_4__t9_funct_4
1. ! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_funct_1(X3) & v1_relat_1(X3)) => (r1_tarski(X2,X3) => r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),k5_relat_1(k6_relat_1(X0,X3),X1))))) [input t9_funct_4]
2. ~! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_funct_1(X3) & v1_relat_1(X3)) => (r1_tarski(X2,X3) => r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),k5_relat_1(k6_relat_1(X0,X3),X1))))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
7. ! [X0,X1] : (v1_relat_1(X1) => v1_relat_1(k6_relat_1(X0,X1))) [input dt_k6_relat_1]
16. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2))))) [input t101_relat_1]
17. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
18. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0))))) [input t76_relat_1]
19. ~! [X0,X1,X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_funct_1(X3) & v1_relat_1(X3)) => (r1_tarski(X2,X3) => r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),k5_relat_1(k6_relat_1(X0,X3),X1))))) [flattening 2]
20. ! [X0,X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0))))) [flattening 18]
21. ! [X0,X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2))))) [flattening 16]
22. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 17]
24. ? [X0,X1,X2] : (? [X3] : ((~r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),k5_relat_1(k6_relat_1(X0,X3),X1)) & r1_tarski(X2,X3)) & (v1_funct_1(X3) & v1_relat_1(X3))) & (v1_funct_1(X2) & v1_relat_1(X2))) [ennf transformation 19]
25. ? [X0,X1,X2] : (? [X3] : (~r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),k5_relat_1(k6_relat_1(X0,X3),X1)) & r1_tarski(X2,X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) [flattening 24]
26. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
29. ! [X0,X1] : (! [X2] : ((r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [ennf transformation 20]
30. ! [X0,X1] : (! [X2] : (r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [flattening 29]
31. ! [X0,X1] : (! [X2] : ((r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2)) | ~r1_tarski(X1,X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [ennf transformation 21]
32. ! [X0,X1] : (! [X2] : (r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [flattening 31]
34. ! [X0,X1] : (v1_relat_1(k6_relat_1(X0,X1)) | ~v1_relat_1(X1)) [ennf transformation 7]
41. ? [X0,X1,X2] : (? [X3] : (~r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),k5_relat_1(k6_relat_1(X0,X3),X1)) & r1_tarski(X2,X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) => (? [X3] : (~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),k5_relat_1(k6_relat_1(sK0,X3),sK1)) & r1_tarski(sK2,X3) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
42. ? [X3] : (~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),k5_relat_1(k6_relat_1(sK0,X3),sK1)) & r1_tarski(sK2,X3) & v1_funct_1(X3) & v1_relat_1(X3)) => (~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),k5_relat_1(k6_relat_1(sK0,sK3),sK1)) & r1_tarski(sK2,sK3) & v1_funct_1(sK3) & v1_relat_1(sK3)) [choice axiom]
43. (~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),k5_relat_1(k6_relat_1(sK0,sK3),sK1)) & r1_tarski(sK2,sK3) & v1_funct_1(sK3) & v1_relat_1(sK3)) & v1_funct_1(sK2) & v1_relat_1(sK2) [skolemisation 25,42,41]
44. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 22]
51. v1_relat_1(sK3) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
53. r1_tarski(sK2,sK3) [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. ~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),k5_relat_1(k6_relat_1(sK0,sK3),sK1)) [cnf transformation 43] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
57. r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 30] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:1,thDist:-1}
58. r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 32] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:1,thDist:-1}
60. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 44] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
64. v1_relat_1(k6_relat_1(X0,X1)) | ~v1_relat_1(X1) [cnf transformation 34] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
74. 1 <=> r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),k5_relat_1(k6_relat_1(sK0,sK3),sK1)) [avatar definition]
76. ~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),k5_relat_1(k6_relat_1(sK0,sK3),sK1)) <- (~1) [avatar component clause 74] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
77. ~1 [avatar split clause 54,74]
79. 2 <=> r1_tarski(sK2,sK3) [avatar definition]
81. r1_tarski(sK2,sK3) <- (2) [avatar component clause 79] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. 2 [avatar split clause 53,79]
89. 4 <=> v1_relat_1(sK3) [avatar definition]
91. v1_relat_1(sK3) <- (4) [avatar component clause 89] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. 4 [avatar split clause 51,89]
116. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 55,60] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
131. r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2) [subsumption resolution 57,116] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
132. ~r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK0,sK3)) | ~v1_relat_1(k6_relat_1(sK0,sK3)) <- (~1) [resolution 131,76] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
138. 9 <=> v1_relat_1(k6_relat_1(sK0,sK3)) [avatar definition]
140. ~v1_relat_1(k6_relat_1(sK0,sK3)) <- (~9) [avatar component clause 138] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
142. 10 <=> r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK0,sK3)) [avatar definition]
144. ~r1_tarski(k6_relat_1(sK0,sK2),k6_relat_1(sK0,sK3)) <- (~10) [avatar component clause 142] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
145. ~9 | ~10 | 1 [avatar split clause 132,74,142,138]
149. r1_tarski(k6_relat_1(X0,X1),k6_relat_1(X0,X2)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2) [subsumption resolution 58,116] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
158. ~v1_relat_1(sK3) <- (~9) [resolution 140,64] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
160. $false <- (4, ~9) [subsumption resolution 158,91] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
161. ~4 | 9 [avatar contradiction clause 160]
163. ~r1_tarski(sK2,sK3) | ~v1_relat_1(sK3) <- (~10) [resolution 144,149] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
165. ~v1_relat_1(sK3) <- (2, ~10) [subsumption resolution 163,81] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
166. $false <- (2, 4, ~10) [subsumption resolution 165,91] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
167. ~2 | ~4 | 10 [avatar contradiction clause 166]
168. $false [avatar sat refutation 77,82,92,145,161,167] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:66,thDist:-66}
% SZS output end Proof for funct_4__t9_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 77
% Active clauses: 30
% Passive clauses: 42
% Final active clauses: 29
% Final passive clauses: 8
% Main loop iterations started: 34
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 19
% Unit resulting resolution: 2
% Forward superposition: 4
% Backward superposition: 4
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 8
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

