i: [1,129,0,2,1,1,0,0] t11_membered
a: [3,170,0,2,1,129]
s: 170
i: [1,128,0,2,1,1,0,0] t11_membered
a: [3,175,0,2,1,128]
s: 175
i: [1,130,0,5,2,1,0,0] t11_membered
s: 130
i: [1,131,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,180,0,2,1,131]
s: 180
i: [1,155,0,2,1,0,0,1] rc2_int_1
a: [3,190,0,2,1,155]
s: 190
i: [1,156,0,2,1,0,0,2] rc4_ordinal1
a: [3,195,0,2,1,156]
s: 195
i: [1,157,0,2,1,0,0,255] rc1_ordinal1
a: [3,200,0,2,1,157]
s: 200
i: [1,161,0,2,1,0,0,255] rc2_ordinal1
a: [3,205,0,2,1,161]
s: 205
i: [1,160,0,2,1,0,0,255] rc2_ordinal1
a: [3,210,0,2,1,160]
s: 210
i: [1,159,0,2,1,0,0,255] rc2_ordinal1
a: [3,215,0,2,1,159]
s: 215
i: [1,158,0,2,1,0,0,255] rc2_ordinal1
a: [3,220,0,2,1,158]
s: 220
i: [1,163,0,2,1,0,0,2] rc5_ordinal1
a: [3,225,0,2,1,163]
s: 225
i: [1,162,0,2,1,0,0,2] rc5_ordinal1
a: [3,230,0,2,1,162]
s: 230
i: [1,164,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,235,0,2,1,164]
s: 235
i: [1,165,0,2,1,0,0,255] rc4_xreal_0
a: [3,240,0,2,1,165]
s: 240
i: [1,166,0,2,1,0,0,255] rc3_xreal_0
a: [3,245,0,2,1,166]
s: 245
i: [1,132,0,3,1,0,0,255] d2_xboole_0
a: [3,185,0,3,1,132]
s: 185
r: 180
d: [2,247,0,2,1,2,23,180,185]
a: [3,251,0,2,1,247]
s: 251
i: [1,134,0,4,2,0,0,2] cc7_ordinal1
s: 134
i: [1,135,0,4,2,0,0,255] cc3_ordinal1
s: 135
i: [1,137,0,4,2,0,0,1] cc2_int_1
s: 137
d: [2,253,1,2,1,1,40,134,240]
a: [3,258,1,2,1,253]
s: 258
i: [1,138,0,4,2,0,0,255] cc6_ordinal1
s: 138
d: [2,254,1,2,1,1,40,134,251]
a: [3,263,1,2,1,254]
s: 263
i: [1,140,0,4,2,0,0,255] cc3_nat_1
s: 140
i: [1,141,0,4,2,0,0,255] cc1_nat_1
s: 141
d: [2,265,1,2,1,1,40,135,240]
a: [3,270,1,2,1,265]
s: 270
i: [1,143,0,4,2,0,0,255] cc1_ordinal1
s: 143
d: [2,266,1,2,1,1,40,135,251]
a: [3,275,1,2,1,266]
s: 275
i: [1,144,0,4,2,0,0,255] cc1_ordinal1
s: 144
i: [1,153,0,4,1,0,0,2] existence_m1_subset_1
s: 153
d: [2,277,1,2,1,1,40,137,195]
a: [3,282,1,2,1,277]
s: 282
i: [1,136,0,5,2,0,0,255] t6_boole
s: 136
d: [2,278,1,2,1,1,40,137,225]
a: [3,287,1,2,1,278]
s: 287
i: [1,147,0,5,2,0,0,1] d1_xboole_0
s: 147
i: [1,152,0,5,2,0,0,1] t7_boole
s: 152
d: [2,295,1,2,1,1,40,138,195]
a: [3,301,1,2,1,295]
s: 301
i: [1,146,0,6,3,0,0,255] cc2_ordinal1
s: 146
d: [2,296,1,2,1,1,40,138,225]
a: [3,306,1,2,1,296]
s: 306
i: [1,148,0,6,2,0,0,1] d1_xboole_0
s: 148
i: [1,149,0,6,2,0,0,1] t1_subset
s: 149
d: [2,390,1,5,2,0,40,148,130]
d: [2,392,1,3,1,1,21,390,175]
a: [3,396,1,3,1,392]
s: 396
i: [1,154,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 154
d: [2,317,1,2,1,1,40,140,245]
a: [3,321,1,2,1,317]
s: 321
i: [1,133,0,7,3,0,0,1] cc16_membered
s: 133
i: [1,145,0,7,3,0,0,255] cc5_ordinal1
s: 145
d: [2,323,1,2,1,1,40,143,200]
a: [3,329,1,2,1,323]
s: 329
i: [1,151,0,7,3,0,0,255] t8_boole
s: 151
d: [2,343,1,2,1,1,40,144,200]
a: [3,354,1,2,1,343]
s: 354
i: [1,150,0,8,3,0,0,1] t2_subset
s: 150
d: [2,400,1,5,2,0,40,133,153]
s: 400
d: [2,289,2,2,1,1,40,258,137]
a: [3,293,2,2,1,289]
s: 293
d: [2,401,1,5,2,0,40,145,153]
s: 401
d: [2,310,2,2,1,1,40,263,137]
a: [3,315,2,2,1,310]
s: 315
d: [2,398,1,6,2,0,40,149,148]
s: 398
d: [2,399,1,6,2,0,40,154,148]
s: 399
d: [2,325,2,2,1,1,40,143,270]
a: [3,334,2,2,1,325]
s: 334
d: [2,403,1,6,2,0,40,150,153]
s: 403
d: [2,411,2,5,2,0,40,403,130]
d: [2,415,2,3,1,1,21,411,175]
a: [3,419,2,3,1,415]
s: 419
d: [2,421,3,0,0,2,52,170,419,400]
e: 421
d: [2,422,3,0,0,2,52,170,153,419,133]
e: 422
d: [2,423,3,2,1,1,40,419,400]
d: [2,426,3,0,0,2,21,423,170]
e: 426
f: 426,415,128,129
% Refutation found. Thanks to Tanya!
% SZS status Theorem for membered__t11_membered
% SZS output start Proof for membered__t11_membered
1. ! [X0] : ((v5_membered(X0) & ~v1_xboole_0(X0)) => ? [X1] : (r2_hidden(X1,X0) & v1_int_1(X1))) [input t11_membered]
2. ~! [X0] : ((v5_membered(X0) & ~v1_xboole_0(X0)) => ? [X1] : (r2_hidden(X1,X0) & v1_int_1(X1))) [negated conjecture 1]
8. ! [X0] : (v5_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_int_1(X1))) [input cc16_membered]
36. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
49. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
54. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 49]
84. ? [X0] : (! [X1] : (~r2_hidden(X1,X0) | ~v1_int_1(X1)) & (v5_membered(X0) & ~v1_xboole_0(X0))) [ennf transformation 2]
85. ? [X0] : (! [X1] : (~r2_hidden(X1,X0) | ~v1_int_1(X1)) & v5_membered(X0) & ~v1_xboole_0(X0)) [flattening 84]
86. ! [X0] : (! [X1] : (v1_int_1(X1) | ~m1_subset_1(X1,X0)) | ~v5_membered(X0)) [ennf transformation 8]
99. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 54]
100. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 99]
104. ? [X0] : (! [X1] : (~r2_hidden(X1,X0) | ~v1_int_1(X1)) & v5_membered(X0) & ~v1_xboole_0(X0)) => (! [X1] : (~r2_hidden(X1,sK0) | ~v1_int_1(X1)) & v5_membered(sK0) & ~v1_xboole_0(sK0)) [choice axiom]
105. ! [X1] : (~r2_hidden(X1,sK0) | ~v1_int_1(X1)) & v5_membered(sK0) & ~v1_xboole_0(sK0) [skolemisation 85,104]
110. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK2(X0),X0)) [choice axiom]
111. ! [X0] : m1_subset_1(sK2(X0),X0) [skolemisation 36,110]
128. ~v1_xboole_0(sK0) [cnf transformation 105] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
129. v5_membered(sK0) [cnf transformation 105] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
130. ~r2_hidden(X1,sK0) | ~v1_int_1(X1) [cnf transformation 105] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
133. ~m1_subset_1(X1,X0) | v1_int_1(X1) | ~v5_membered(X0) [cnf transformation 86] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
150. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 100] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
153. m1_subset_1(sK2(X0),X0) [cnf transformation 111] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
168. 1 <=> v5_membered(sK0) [avatar definition]
170. v5_membered(sK0) <- (1) [avatar component clause 168] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
171. 1 [avatar split clause 129,168]
173. 2 <=> v1_xboole_0(sK0) [avatar definition]
175. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 173] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
176. ~2 [avatar split clause 128,173]
400. v1_int_1(sK2(X0)) | ~v5_membered(X0) [resolution 133,153] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
403. r2_hidden(sK2(X0),X0) | v1_xboole_0(X0) [resolution 150,153] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
411. v1_xboole_0(sK0) | ~v1_int_1(sK2(sK0)) [resolution 403,130] {a:2,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
415. ~v1_int_1(sK2(sK0)) <- (~2) [subsumption resolution 411,175] {a:2,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
417. 40 <=> v1_int_1(sK2(sK0)) [avatar definition]
419. ~v1_int_1(sK2(sK0)) <- (~40) [avatar component clause 417] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
420. ~40 | 2 [avatar split clause 415,173,417]
423. ~v5_membered(sK0) <- (~40) [resolution 419,400] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
426. $false <- (1, ~40) [subsumption resolution 423,170] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
427. ~1 | 40 [avatar contradiction clause 426]
428. $false [avatar sat refutation 171,176,420,427] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:212,thDist:-212}
% SZS output end Proof for membered__t11_membered
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 51
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 10
% 
% >>> Saturation
% Initial clauses: 39
% Generated clauses: 135
% Active clauses: 57
% Passive clauses: 73
% Final active clauses: 56
% Final passive clauses: 14
% Main loop iterations started: 59
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 44
% Unit resulting resolution: 2
% Forward superposition: 1
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 47
% SAT solver unit clauses: 16
% SAT solver binary clauses: 28
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% property evaluation: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% other: 0.003 s ( own 0.001 s ) 

