i: [1,34,0,8,3,1,0,0] t41_gate_1
a: [3,55,0,2,1,34]
s: 55
a: [3,60,0,2,1,34]
s: 60
a: [3,51,0,4,1,34]
s: 51
i: [1,36,0,2,1,0,0,255] fc1_xboole_0
a: [3,68,0,2,1,36]
s: 68
i: [1,47,0,2,1,0,0,1] rc2_xboole_0
a: [3,73,0,2,1,47]
s: 73
i: [1,48,0,2,1,0,0,1] rc1_xboole_0
a: [3,78,0,2,1,48]
s: 78
i: [1,37,0,5,2,0,0,255] fc2_gate_1
s: 37
i: [1,38,0,5,2,0,0,255] fc1_gate_1
s: 38
i: [1,45,0,7,1,0,0,255] commutativity_k34_gate_1
s: 45
i: [1,46,0,7,1,0,0,255] commutativity_k34_gate_1
s: 46
i: [1,43,0,9,3,0,0,255] d34_gate_1
s: 43
i: [1,44,0,9,3,0,0,255] d34_gate_1
s: 44
d: [2,88,1,6,3,1,44,51,44]
d: [2,91,1,4,2,2,21,88,68]
d: [2,92,1,2,1,3,21,91,60]
d: [2,93,1,0,0,4,21,92,55]
e: 93
r: 51
r: 60
a: [3,59,0,2,1,34]
s: 59
a: [3,52,0,4,1,34]
s: 52
i: [1,40,0,10,3,0,0,255] d34_gate_1
s: 40
i: [1,41,0,10,3,0,0,255] d34_gate_1
s: 41
d: [2,118,1,7,3,1,44,52,41]
d: [2,123,1,5,2,2,21,118,59]
d: [2,124,1,3,1,3,21,123,55]
a: [3,128,1,3,1,124]
s: 128
d: [2,130,2,0,0,2,52,68,128,38]
e: 130
d: [2,131,2,2,1,1,40,128,38]
d: [2,133,2,0,0,2,21,131,68]
e: 133
r: 55
r: 59
r: 128
r: 52
a: [3,56,0,2,1,34]
s: 56
s: 59
s: 51
d: [2,140,1,6,3,1,44,51,43]
d: [2,141,1,4,2,2,21,140,68]
d: [2,142,1,2,1,3,21,141,59]
d: [2,143,1,0,0,4,21,142,56]
e: 143
i: [1,33,0,8,3,1,0,0] t41_gate_1
i: [1,31,0,8,3,1,0,0] t41_gate_1
i: [1,30,0,8,3,1,0,0] t41_gate_1
d: [2,105,1,7,3,1,44,52,40]
f: 143,118,105,133,93,36,30,31,33,34
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t41_gate_1
% SZS output start Proof for gate_1__t41_gate_1
1. ! [X0] : ! [X1] : (~v1_xboole_0(k34_gate_1(X0,X1)) <=> (~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0)))) [input t41_gate_1]
2. ~! [X0] : ! [X1] : (~v1_xboole_0(k34_gate_1(X0,X1)) <=> (~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0)))) [negated conjecture 1]
10. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
11. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
15. ! [X0] : ! [X1] : ((~(~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0))) => k34_gate_1(X0,X1) = k1_xboole_0) & ~(k34_gate_1(X0,X1) != k1_gate_1(k1_xboole_0) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0)))) [input d34_gate_1]
16. ~! [X0,X1] : (~v1_xboole_0(k34_gate_1(X0,X1)) <=> (~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0)))) [flattening 2]
17. ! [X0,X1] : ((~(~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0))) => k34_gate_1(X0,X1) = k1_xboole_0) & ~(k34_gate_1(X0,X1) != k1_gate_1(k1_xboole_0) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & v1_xboole_0(X0)))) [flattening 15]
18. ? [X0,X1] : (~v1_xboole_0(k34_gate_1(X0,X1)) <~> ((v1_xboole_0(X1) | v1_xboole_0(X0)) & (~v1_xboole_0(X1) | ~v1_xboole_0(X0)))) [ennf transformation 16]
20. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 10]
21. ! [X0,X1] : ((k34_gate_1(X0,X1) = k1_xboole_0 | ((v1_xboole_0(X1) | v1_xboole_0(X0)) & (~v1_xboole_0(X1) | ~v1_xboole_0(X0)))) & (k34_gate_1(X0,X1) = k1_gate_1(k1_xboole_0) | (~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | (v1_xboole_0(X1) & v1_xboole_0(X0)))) [ennf transformation 17]
22. ? [X0,X1] : ((((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | (v1_xboole_0(X1) & v1_xboole_0(X0))) | v1_xboole_0(k34_gate_1(X0,X1))) & (((v1_xboole_0(X1) | v1_xboole_0(X0)) & (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) | ~v1_xboole_0(k34_gate_1(X0,X1)))) [nnf transformation 18]
23. ? [X0,X1] : (((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | (v1_xboole_0(X1) & v1_xboole_0(X0)) | v1_xboole_0(k34_gate_1(X0,X1))) & (((v1_xboole_0(X1) | v1_xboole_0(X0)) & (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) | ~v1_xboole_0(k34_gate_1(X0,X1)))) [flattening 22]
24. ? [X0,X1] : (((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | (v1_xboole_0(X1) & v1_xboole_0(X0)) | v1_xboole_0(k34_gate_1(X0,X1))) & (((v1_xboole_0(X1) | v1_xboole_0(X0)) & (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) | ~v1_xboole_0(k34_gate_1(X0,X1)))) => (((~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | (v1_xboole_0(sK1) & v1_xboole_0(sK0)) | v1_xboole_0(k34_gate_1(sK0,sK1))) & (((v1_xboole_0(sK1) | v1_xboole_0(sK0)) & (~v1_xboole_0(sK1) | ~v1_xboole_0(sK0))) | ~v1_xboole_0(k34_gate_1(sK0,sK1)))) [choice axiom]
25. ((~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | (v1_xboole_0(sK1) & v1_xboole_0(sK0)) | v1_xboole_0(k34_gate_1(sK0,sK1))) & (((v1_xboole_0(sK1) | v1_xboole_0(sK0)) & (~v1_xboole_0(sK1) | ~v1_xboole_0(sK0))) | ~v1_xboole_0(k34_gate_1(sK0,sK1))) [skolemisation 23,24]
30. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | ~v1_xboole_0(k34_gate_1(sK0,sK1)) [cnf transformation 25] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. v1_xboole_0(sK1) | v1_xboole_0(sK0) | ~v1_xboole_0(k34_gate_1(sK0,sK1)) [cnf transformation 25] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. ~v1_xboole_0(sK0) | v1_xboole_0(sK1) | v1_xboole_0(k34_gate_1(sK0,sK1)) [cnf transformation 25] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. ~v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k34_gate_1(sK0,sK1)) [cnf transformation 25] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. v1_xboole_0(k1_xboole_0) [cnf transformation 11] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
38. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 20] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. k34_gate_1(X0,X1) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X0) | v1_xboole_0(X1) [cnf transformation 21] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
41. k34_gate_1(X0,X1) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 21] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. k34_gate_1(X0,X1) = k1_xboole_0 | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 21] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. k34_gate_1(X0,X1) = k1_xboole_0 | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 21] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
50. 1 <=> v1_xboole_0(k34_gate_1(sK0,sK1)) [avatar definition]
51. ~v1_xboole_0(k34_gate_1(sK0,sK1)) <- (~1) [avatar component clause 50] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
52. v1_xboole_0(k34_gate_1(sK0,sK1)) <- (1) [avatar component clause 50] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
54. 2 <=> v1_xboole_0(sK0) [avatar definition]
55. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 54] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. v1_xboole_0(sK0) <- (2) [avatar component clause 54] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
58. 3 <=> v1_xboole_0(sK1) [avatar definition]
59. v1_xboole_0(sK1) <- (3) [avatar component clause 58] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. ~v1_xboole_0(sK1) <- (~3) [avatar component clause 58] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
61. 1 | 2 | ~3 [avatar split clause 34,58,54,50]
62. 1 | 3 | ~2 [avatar split clause 33,54,58,50]
63. ~1 | 2 | 3 [avatar split clause 31,58,54,50]
64. ~1 | ~2 | ~3 [avatar split clause 30,58,54,50]
66. 4 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
68. v1_xboole_0(k1_xboole_0) <- (4) [avatar component clause 66] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
69. 4 [avatar split clause 36,66]
88. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1) [superposition 51,44] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1, 4) [subsumption resolution 88,68] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. v1_xboole_0(sK0) <- (~1, ~3, 4) [subsumption resolution 91,60] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
93. $false <- (~1, ~2, ~3, 4) [subsumption resolution 92,55] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
94. 1 | 2 | 3 | ~4 [avatar contradiction clause 93]
105. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK0) | v1_xboole_0(sK1) <- (1) [superposition 52,40] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1) [superposition 52,41] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. 7 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
128. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (7) [avatar component clause 126] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
131. ~v1_xboole_0(k1_xboole_0) <- (7) [resolution 128,38] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
133. $false <- (4, 7) [subsumption resolution 131,68] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
134. ~4 | ~7 [avatar contradiction clause 133]
135. 3 | ~2 | 7 | ~1 [avatar split clause 105,50,126,54,58]
136. 2 | ~3 | 7 | ~1 [avatar split clause 118,50,126,58,54]
140. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~1) [superposition 51,43] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
141. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~1, 4) [subsumption resolution 140,68] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. ~v1_xboole_0(sK0) <- (~1, 3, 4) [subsumption resolution 141,59] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
143. $false <- (~1, 2, 3, 4) [subsumption resolution 142,56] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 143]
145. $false [avatar sat refutation 61,62,63,64,69,94,134,135,136,144] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:64,thDist:-64}
% SZS output end Proof for gate_1__t41_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 89
% Active clauses: 19
% Passive clauses: 26
% Final active clauses: 13
% Final passive clauses: 5
% Main loop iterations started: 20
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Fw subsumption resolutions: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 20
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 18
% Backward superposition: 16
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 18
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 3
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward demodulation: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

