i: [1,28,0,2,1,1,0,0] t10_zf_refle
a: [3,49,0,2,1,28]
s: 49
i: [1,27,0,2,1,1,0,0] t10_zf_refle
a: [3,54,0,2,1,27]
s: 54
i: [1,26,0,2,1,1,0,0] t10_zf_refle
a: [3,59,0,2,1,26]
s: 59
i: [1,25,0,2,1,1,0,0] t10_zf_refle
a: [3,64,0,2,1,25]
s: 64
i: [1,29,0,3,1,1,0,0] t10_zf_refle
a: [3,44,0,3,1,29]
s: 44
i: [1,34,0,2,1,0,0,1] rc1_ordinal1
a: [3,69,0,2,1,34]
s: 69
i: [1,38,0,2,1,0,0,1] rc2_ordinal2
a: [3,74,0,2,1,38]
s: 74
i: [1,37,0,2,1,0,0,1] rc2_ordinal2
a: [3,79,0,2,1,37]
s: 79
i: [1,36,0,2,1,0,0,1] rc2_ordinal2
a: [3,84,0,2,1,36]
s: 84
i: [1,35,0,2,1,0,0,1] rc2_ordinal2
a: [3,89,0,2,1,35]
s: 89
i: [1,40,0,2,1,0,0,1] rc1_funct_1
a: [3,94,0,2,1,40]
s: 94
i: [1,39,0,2,1,0,0,1] rc1_funct_1
a: [3,99,0,2,1,39]
s: 99
i: [1,30,0,4,2,0,0,255] cc1_ordinal1
s: 30
d: [2,101,1,2,1,1,40,30,69]
a: [3,105,1,2,1,101]
s: 105
i: [1,31,0,4,2,0,0,255] cc1_ordinal1
s: 31
i: [1,32,0,6,3,0,0,255] cc2_ordinal1
s: 32
d: [2,107,1,2,1,1,40,31,69]
a: [3,111,1,2,1,107]
s: 111
i: [1,33,0,11,5,0,0,1] fc1_card_5
s: 33
d: [2,114,1,0,0,5,52,64,54,59,49,44,33]
e: 114
d: [2,115,1,8,4,1,40,33,44]
d: [2,119,1,6,3,2,21,115,49]
d: [2,120,1,4,2,3,21,119,59]
d: [2,121,1,2,1,4,21,120,54]
d: [2,122,1,0,0,5,21,121,64]
e: 122
f: 122,25,26,27,28,29
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zf_refle__t10_zf_refle
% SZS output start Proof for zf_refle__t10_zf_refle
1. ! [X0] : ((v1_ordinal2(X0) & v1_funct_1(X0) & v5_ordinal1(X0) & v1_relat_1(X0)) => v3_ordinal1(k3_card_3(X0))) [input t10_zf_refle]
2. ~! [X0] : ((v1_ordinal2(X0) & v1_funct_1(X0) & v5_ordinal1(X0) & v1_relat_1(X0)) => v3_ordinal1(k3_card_3(X0))) [negated conjecture 1]
6. ! [X0] : ((v1_ordinal2(X0) & v5_ordinal1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => v3_ordinal1(k3_card_3(X0))) [input fc1_card_5]
10. ? [X0] : (~v3_ordinal1(k3_card_3(X0)) & (v1_ordinal2(X0) & v1_funct_1(X0) & v5_ordinal1(X0) & v1_relat_1(X0))) [ennf transformation 2]
11. ? [X0] : (~v3_ordinal1(k3_card_3(X0)) & v1_ordinal2(X0) & v1_funct_1(X0) & v5_ordinal1(X0) & v1_relat_1(X0)) [flattening 10]
15. ! [X0] : (v3_ordinal1(k3_card_3(X0)) | (~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 6]
16. ! [X0] : (v3_ordinal1(k3_card_3(X0)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 15]
17. ? [X0] : (~v3_ordinal1(k3_card_3(X0)) & v1_ordinal2(X0) & v1_funct_1(X0) & v5_ordinal1(X0) & v1_relat_1(X0)) => (~v3_ordinal1(k3_card_3(sK0)) & v1_ordinal2(sK0) & v1_funct_1(sK0) & v5_ordinal1(sK0) & v1_relat_1(sK0)) [choice axiom]
18. ~v3_ordinal1(k3_card_3(sK0)) & v1_ordinal2(sK0) & v1_funct_1(sK0) & v5_ordinal1(sK0) & v1_relat_1(sK0) [skolemisation 11,17]
25. v1_relat_1(sK0) [cnf transformation 18] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
26. v5_ordinal1(sK0) [cnf transformation 18] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. v1_funct_1(sK0) [cnf transformation 18] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. v1_ordinal2(sK0) [cnf transformation 18] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~v3_ordinal1(k3_card_3(sK0)) [cnf transformation 18] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. v3_ordinal1(k3_card_3(X0)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 16] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. 1 <=> v3_ordinal1(k3_card_3(sK0)) [avatar definition]
44. ~v3_ordinal1(k3_card_3(sK0)) <- (~1) [avatar component clause 42] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
45. ~1 [avatar split clause 29,42]
47. 2 <=> v1_ordinal2(sK0) [avatar definition]
49. v1_ordinal2(sK0) <- (2) [avatar component clause 47] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
50. 2 [avatar split clause 28,47]
52. 3 <=> v1_funct_1(sK0) [avatar definition]
54. v1_funct_1(sK0) <- (3) [avatar component clause 52] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
55. 3 [avatar split clause 27,52]
57. 4 <=> v5_ordinal1(sK0) [avatar definition]
59. v5_ordinal1(sK0) <- (4) [avatar component clause 57] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 4 [avatar split clause 26,57]
62. 5 <=> v1_relat_1(sK0) [avatar definition]
64. v1_relat_1(sK0) <- (5) [avatar component clause 62] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. 5 [avatar split clause 25,62]
115. ~v1_ordinal2(sK0) | ~v5_ordinal1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [resolution 33,44] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
119. ~v5_ordinal1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 115,49] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
120. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 4) [subsumption resolution 119,59] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
121. ~v1_relat_1(sK0) <- (~1, 2, 3, 4) [subsumption resolution 120,54] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
122. $false <- (~1, 2, 3, 4, 5) [subsumption resolution 121,64] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
123. 1 | ~2 | ~3 | ~4 | ~5 [avatar contradiction clause 122]
124. $false [avatar sat refutation 45,50,55,60,65,123] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:60,thDist:-60}
% SZS output end Proof for zf_refle__t10_zf_refle
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 16
% Generated clauses: 41
% Active clauses: 18
% Passive clauses: 18
% Final active clauses: 18
% Main loop iterations started: 18
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 4
% 
% >>> Deletion Inferences
% Forward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 6
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 12
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

