i: [1,34,0,5,1,1,0,0] t119_zfmisc_1
a: [3,52,0,5,1,34]
s: 52
i: [1,33,0,5,1,1,0,0] t119_zfmisc_1
a: [3,57,0,5,1,33]
s: 57
i: [1,35,0,11,1,1,0,0] t119_zfmisc_1
a: [3,47,0,11,1,35]
s: 47
i: [1,41,0,3,1,0,0,1] reflexivity_r1_tarski
s: 41
i: [1,36,0,5,1,0,0,1] t7_xboole_1
s: 36
i: [1,42,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 42
i: [1,43,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 43
i: [1,37,0,9,3,0,0,1] t1_xboole_1
s: 37
d: [2,62,1,8,2,1,40,37,52]
s: 62
i: [1,38,0,11,3,0,0,1] t8_xboole_1
s: 38
i: [1,39,0,13,3,0,0,1] t96_zfmisc_1
s: 39
d: [2,63,1,8,2,1,40,37,57]
s: 63
i: [1,40,0,13,3,0,0,1] t13_xboole_1
s: 40
d: [2,67,1,18,2,1,40,38,47]
a: [3,76,1,9,1,67]
s: 76
d: [2,60,1,5,1,0,44,36,43]
s: 60
d: [2,61,1,5,1,0,44,36,43]
s: 61
d: [2,66,2,11,3,1,40,62,37]
s: 66
d: [2,65,1,8,2,0,40,37,36]
s: 65
d: [2,83,2,11,3,1,40,63,37]
s: 83
d: [2,85,1,11,3,0,44,40,42]
s: 85
d: [2,86,1,13,3,0,44,40,43]
s: 86
d: [2,95,3,14,4,1,40,66,37]
s: 95
d: [2,87,1,13,3,0,44,40,43]
s: 87
d: [2,100,3,14,4,1,40,83,37]
s: 100
d: [2,89,1,13,3,0,44,40,43]
s: 89
d: [2,90,1,13,3,0,44,40,43]
s: 90
d: [2,108,4,17,5,1,40,95,37]
s: 108
d: [2,68,1,14,4,0,40,38,37]
s: 68
d: [2,109,4,17,5,1,40,100,37]
s: 109
d: [2,82,1,16,4,0,40,39,37]
s: 82
d: [2,123,2,0,0,2,52,36,36,76,57,82]
e: 123
r: 76
a: [3,75,1,9,1,67]
s: 75
a: [3,80,1,9,1,67]
s: 80
d: [2,140,2,0,0,2,52,60,52,60,80,82]
e: 140
f: 140,123,67,33,34,35
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t119_zfmisc_1
% SZS output start Proof for zfmisc_1__t119_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ((r1_tarski(X1,k2_zfmisc_1(X4,X5)) & r1_tarski(X0,k2_zfmisc_1(X2,X3))) => r1_tarski(k2_xboole_0(X0,X1),k2_zfmisc_1(k2_xboole_0(X2,X4),k2_xboole_0(X3,X5)))) [input t119_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ((r1_tarski(X1,k2_zfmisc_1(X4,X5)) & r1_tarski(X0,k2_zfmisc_1(X2,X3))) => r1_tarski(k2_xboole_0(X0,X1),k2_zfmisc_1(k2_xboole_0(X2,X4),k2_xboole_0(X3,X5)))) [negated conjecture 1]
3. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
9. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
10. ! [X0] : ! [X1] : r1_tarski(X0,k2_xboole_0(X0,X1)) [input t7_xboole_1]
11. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [input t8_xboole_1]
12. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3))) [input t96_zfmisc_1]
13. ~! [X0,X1,X2,X3,X4,X5] : ((r1_tarski(X1,k2_zfmisc_1(X4,X5)) & r1_tarski(X0,k2_zfmisc_1(X2,X3))) => r1_tarski(k2_xboole_0(X0,X1),k2_zfmisc_1(k2_xboole_0(X2,X4),k2_xboole_0(X3,X5)))) [flattening 2]
14. ! [X0,X1] : r1_tarski(X0,k2_xboole_0(X0,X1)) [flattening 10]
15. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 9]
16. ! [X0,X1,X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [flattening 11]
17. ! [X0,X1,X2,X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3))) [flattening 12]
21. ? [X0,X1,X2,X3,X4,X5] : (~r1_tarski(k2_xboole_0(X0,X1),k2_zfmisc_1(k2_xboole_0(X2,X4),k2_xboole_0(X3,X5))) & (r1_tarski(X1,k2_zfmisc_1(X4,X5)) & r1_tarski(X0,k2_zfmisc_1(X2,X3)))) [ennf transformation 13]
22. ? [X0,X1,X2,X3,X4,X5] : (~r1_tarski(k2_xboole_0(X0,X1),k2_zfmisc_1(k2_xboole_0(X2,X4),k2_xboole_0(X3,X5))) & r1_tarski(X1,k2_zfmisc_1(X4,X5)) & r1_tarski(X0,k2_zfmisc_1(X2,X3))) [flattening 21]
23. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 15]
24. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 23]
25. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | (~r1_tarski(X2,X1) | ~r1_tarski(X0,X1))) [ennf transformation 16]
26. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1)) [flattening 25]
27. ! [X0,X1,X2,X3] : (r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | (~r1_tarski(X2,X3) | ~r1_tarski(X0,X1))) [ennf transformation 17]
28. ! [X0,X1,X2,X3] : (r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | ~r1_tarski(X2,X3) | ~r1_tarski(X0,X1)) [flattening 27]
31. ? [X0,X1,X2,X3,X4,X5] : (~r1_tarski(k2_xboole_0(X0,X1),k2_zfmisc_1(k2_xboole_0(X2,X4),k2_xboole_0(X3,X5))) & r1_tarski(X1,k2_zfmisc_1(X4,X5)) & r1_tarski(X0,k2_zfmisc_1(X2,X3))) => (~r1_tarski(k2_xboole_0(sK0,sK1),k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) & r1_tarski(sK1,k2_zfmisc_1(sK4,sK5)) & r1_tarski(sK0,k2_zfmisc_1(sK2,sK3))) [choice axiom]
32. ~r1_tarski(k2_xboole_0(sK0,sK1),k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) & r1_tarski(sK1,k2_zfmisc_1(sK4,sK5)) & r1_tarski(sK0,k2_zfmisc_1(sK2,sK3)) [skolemisation 22,31]
33. r1_tarski(sK0,k2_zfmisc_1(sK2,sK3)) [cnf transformation 32] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. r1_tarski(sK1,k2_zfmisc_1(sK4,sK5)) [cnf transformation 32] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. ~r1_tarski(k2_xboole_0(sK0,sK1),k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) [cnf transformation 32] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. r1_tarski(X0,k2_xboole_0(X0,X1)) [cnf transformation 14] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 24] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
38. r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1) [cnf transformation 26] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
39. r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | ~r1_tarski(X2,X3) | ~r1_tarski(X0,X1) [cnf transformation 28] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
45. 1 <=> r1_tarski(k2_xboole_0(sK0,sK1),k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) [avatar definition]
47. ~r1_tarski(k2_xboole_0(sK0,sK1),k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) <- (~1) [avatar component clause 45] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
48. ~1 [avatar split clause 35,45]
50. 2 <=> r1_tarski(sK1,k2_zfmisc_1(sK4,sK5)) [avatar definition]
52. r1_tarski(sK1,k2_zfmisc_1(sK4,sK5)) <- (2) [avatar component clause 50] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
53. 2 [avatar split clause 34,50]
55. 3 <=> r1_tarski(sK0,k2_zfmisc_1(sK2,sK3)) [avatar definition]
57. r1_tarski(sK0,k2_zfmisc_1(sK2,sK3)) <- (3) [avatar component clause 55] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
58. 3 [avatar split clause 33,55]
60. r1_tarski(X0,k2_xboole_0(X1,X0)) [superposition 36,43] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
67. ~r1_tarski(sK1,k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) | ~r1_tarski(sK0,k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) <- (~1) [resolution 38,47] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
74. 4 <=> r1_tarski(sK0,k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) [avatar definition]
76. ~r1_tarski(sK0,k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) <- (~4) [avatar component clause 74] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
78. 5 <=> r1_tarski(sK1,k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) [avatar definition]
80. ~r1_tarski(sK1,k2_zfmisc_1(k2_xboole_0(sK2,sK4),k2_xboole_0(sK3,sK5))) <- (~5) [avatar component clause 78] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
81. ~4 | ~5 | 1 [avatar split clause 67,45,78,74]
82. ~r1_tarski(X4,k2_zfmisc_1(X2,X0)) | ~r1_tarski(X2,X3) | r1_tarski(X4,k2_zfmisc_1(X3,X1)) | ~r1_tarski(X0,X1) [resolution 39,37] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
123. $false <- (3, ~4) [unit resulting resolution 36,36,76,57,82] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
137. ~3 | 4 [avatar contradiction clause 123]
140. $false <- (2, ~5) [unit resulting resolution 60,52,60,80,82] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
141. ~2 | 5 [avatar contradiction clause 140]
142. $false [avatar sat refutation 48,53,58,81,137,141] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:35,thDist:-35}
% SZS output end Proof for zfmisc_1__t119_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 11
% Generated clauses: 101
% Active clauses: 27
% Passive clauses: 57
% Final active clauses: 26
% Final passive clauses: 25
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 24
% 
% >>> Generating Inferences
% Binary resolution: 41
% Unit resulting resolution: 2
% Forward superposition: 24
% Backward superposition: 3
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 8
% SAT solver unit clauses: 3
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

