i: [1,88,0,2,1,1,0,0] t2_domain_1
a: [3,154,0,2,1,88]
s: 154
i: [1,87,0,2,1,1,0,0] t2_domain_1
a: [3,159,0,2,1,87]
s: 159
i: [1,93,0,3,1,1,0,0] t2_domain_1
a: [3,129,0,3,1,93]
s: 129
i: [1,92,0,5,1,1,0,0] t2_domain_1
a: [3,134,0,5,1,92]
s: 134
i: [1,91,0,5,1,1,0,0] t2_domain_1
a: [3,139,0,5,1,91]
s: 139
i: [1,90,0,5,1,1,0,0] t2_domain_1
a: [3,144,0,5,1,90]
s: 144
i: [1,89,0,5,1,1,0,0] t2_domain_1
a: [3,149,0,5,1,89]
s: 149
i: [1,94,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,164,0,2,1,94]
s: 164
i: [1,95,0,2,1,0,0,255] fc1_xboole_0
a: [3,169,0,2,1,95]
s: 169
i: [1,120,0,2,1,0,0,1] rc2_xboole_0
a: [3,179,0,2,1,120]
s: 179
i: [1,121,0,2,1,0,0,255] rc1_xtuple_0
a: [3,184,0,2,1,121]
s: 184
i: [1,122,0,2,1,0,0,1] rc1_xboole_0
a: [3,189,0,2,1,122]
s: 189
i: [1,124,0,2,1,0,0,2] rc1_relat_1
a: [3,194,0,2,1,124]
s: 194
i: [1,123,0,2,1,0,0,2] rc1_relat_1
a: [3,199,0,2,1,123]
s: 199
i: [1,125,0,2,1,0,0,255] rc2_relat_1
a: [3,204,0,2,1,125]
s: 204
i: [1,97,0,3,1,0,0,255] fc2_xboole_0
s: 97
i: [1,96,0,3,1,0,0,255] d2_xboole_0
a: [3,174,0,3,1,96]
s: 174
r: 164
i: [1,99,0,4,2,0,0,2] cc1_relat_1
s: 99
d: [2,207,1,2,1,1,40,99,189]
a: [3,212,1,2,1,207]
s: 212
i: [1,109,0,4,1,0,0,1] existence_m1_subset_1
s: 109
i: [1,110,0,4,1,0,0,255] fc3_xboole_0
s: 110
d: [2,208,1,2,1,1,40,99,169]
a: [3,217,1,2,1,208]
s: 217
i: [1,111,0,4,1,0,0,255] fc1_xtuple_0
s: 111
i: [1,112,0,4,1,0,0,1] fc6_relat_1
s: 112
i: [1,100,0,5,2,0,0,255] t6_boole
s: 100
i: [1,108,0,5,2,0,0,2] t7_boole
s: 108
i: [1,113,0,5,1,0,0,255] fc5_relat_1
s: 113
i: [1,105,0,6,2,0,0,1] t1_subset
s: 105
i: [1,114,0,6,1,0,0,1] rd1_xtuple_0
s: 114
i: [1,115,0,6,1,0,0,1] rd2_xtuple_0
s: 115
i: [1,117,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 117
i: [1,107,0,7,3,0,0,1] t8_boole
s: 107
i: [1,116,0,7,1,0,0,255] commutativity_k2_tarski
s: 116
d: [2,219,1,7,3,0,44,100,100]
s: 219
i: [1,106,0,8,3,0,0,1] t2_subset
s: 106
i: [1,118,0,8,3,0,0,1] fc10_subset_1
s: 118
d: [2,224,1,9,2,1,40,106,144]
a: [3,234,1,4,1,224]
s: 234
d: [2,241,2,0,0,3,52,159,154,234,118]
e: 241
d: [2,242,2,4,2,1,40,234,118]
d: [2,245,2,2,1,2,21,242,154]
d: [2,246,2,0,0,3,21,245,159]
e: 246
r: 234
i: [1,119,0,8,1,0,0,255] fc7_relat_1
s: 119
a: [3,233,1,4,1,224]
s: 233
i: [1,98,0,9,2,0,0,255] rd3_xtuple_0
s: 98
i: [1,103,0,10,1,0,0,255] d5_tarski
s: 103
d: [2,250,1,9,2,1,44,98,139]
d: [2,258,1,9,2,2,22,250,134]
a: [3,262,1,2,1,258]
s: 262
i: [1,104,0,12,3,0,0,1] t21_mcart_1
s: 104
a: [3,230,1,5,1,224]
s: 230
d: [2,225,1,9,2,1,40,106,149]
a: [3,239,1,5,1,225]
s: 239
d: [2,280,1,12,3,1,44,104,139]
d: [2,288,1,12,3,2,22,280,134]
a: [3,291,1,5,2,288]
s: 291
d: [2,311,2,0,0,2,52,112,230,291]
e: 311
d: [2,312,2,4,1,2,40,291,230]
d: [2,314,2,0,0,2,21,312,112]
e: 314
r: 291
d: [2,295,2,5,1,1,40,230,117]
a: [3,301,2,5,1,295]
s: 301
a: [3,266,1,7,1,258]
s: 266
d: [2,325,2,2,1,1,44,111,266]
d: [2,342,2,0,0,2,21,325,262]
e: 342
r: 262
a: [3,261,1,2,1,258]
s: 261
d: [2,279,1,4,1,0,44,110,103]
s: 279
d: [2,256,1,5,2,0,44,113,98]
s: 256
d: [2,317,2,5,2,1,44,266,98]
d: [2,327,2,2,1,2,21,317,129]
a: [3,331,2,2,1,327]
s: 331
d: [2,226,1,6,2,0,40,106,109]
s: 226
d: [2,347,2,2,1,1,44,279,266]
a: [3,351,2,2,1,347]
s: 351
d: [2,287,1,7,3,0,44,111,104]
s: 287
d: [2,356,3,0,0,2,52,331,112,239,287]
e: 356
d: [2,358,2,6,2,1,40,287,239]
d: [2,361,2,4,1,2,21,358,331]
d: [2,362,2,0,0,2,21,361,112]
e: 362
d: [2,282,1,12,3,1,44,104,134]
d: [2,293,1,12,3,2,22,282,139]
d: [2,319,2,5,2,1,44,98,266]
d: [2,334,2,2,1,2,21,319,129]
f: 362,334,314,293,246,225,224,87,88,89,90,91,92,93
% Refutation found. Thanks to Tanya!
% SZS status Theorem for domain_1__t2_domain_1
% SZS output start Proof for domain_1__t2_domain_1
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,k2_zfmisc_1(X0,X1)) => ! [X3] : (m1_subset_1(X3,k2_zfmisc_1(X0,X1)) => ((k2_xtuple_0(X2) = k2_xtuple_0(X3) & k1_xtuple_0(X2) = k1_xtuple_0(X3)) => X2 = X3))))) [input t2_domain_1]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,k2_zfmisc_1(X0,X1)) => ! [X3] : (m1_subset_1(X3,k2_zfmisc_1(X0,X1)) => ((k2_xtuple_0(X2) = k2_xtuple_0(X3) & k1_xtuple_0(X2) = k1_xtuple_0(X3)) => X2 = X3))))) [negated conjecture 1]
20. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_zfmisc_1(X0,X1))) [input fc10_subset_1]
22. ! [X0,X1] : v1_xtuple_0(k4_tarski(X0,X1)) [input fc1_xtuple_0]
26. ! [X0,X1] : v1_relat_1(k2_zfmisc_1(X0,X1)) [input fc6_relat_1]
35. ! [X0] : (v1_xtuple_0(X0) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0) [input rd3_xtuple_0]
37. ! [X0] : ! [X1] : (v1_relat_1(X1) => (r2_hidden(X0,X1) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0)) [input t21_mcart_1]
38. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
43. ! [X0,X1] : (v1_relat_1(X1) => (r2_hidden(X0,X1) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0)) [flattening 37]
45. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 38]
51. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((X2 != X3 & (k2_xtuple_0(X2) = k2_xtuple_0(X3) & k1_xtuple_0(X2) = k1_xtuple_0(X3))) & m1_subset_1(X3,k2_zfmisc_1(X0,X1))) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [ennf transformation 2]
52. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (X2 != X3 & k2_xtuple_0(X2) = k2_xtuple_0(X3) & k1_xtuple_0(X2) = k1_xtuple_0(X3) & m1_subset_1(X3,k2_zfmisc_1(X0,X1))) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [flattening 51]
53. ! [X0] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0)) [ennf transformation 35]
60. ! [X0,X1] : ((k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~r2_hidden(X0,X1)) | ~v1_relat_1(X1)) [ennf transformation 43]
61. ! [X0,X1] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~r2_hidden(X0,X1) | ~v1_relat_1(X1)) [flattening 60]
63. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 45]
64. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 63]
68. ! [X0,X1] : (~v1_xboole_0(k2_zfmisc_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 20]
69. ! [X0,X1] : (~v1_xboole_0(k2_zfmisc_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 68]
70. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (X2 != X3 & k2_xtuple_0(X2) = k2_xtuple_0(X3) & k1_xtuple_0(X2) = k1_xtuple_0(X3) & m1_subset_1(X3,k2_zfmisc_1(X0,X1))) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (X2 != X3 & k2_xtuple_0(X2) = k2_xtuple_0(X3) & k1_xtuple_0(X2) = k1_xtuple_0(X3) & m1_subset_1(X3,k2_zfmisc_1(sK0,X1))) & m1_subset_1(X2,k2_zfmisc_1(sK0,X1))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
71. ? [X1] : (? [X2] : (? [X3] : (X2 != X3 & k2_xtuple_0(X2) = k2_xtuple_0(X3) & k1_xtuple_0(X2) = k1_xtuple_0(X3) & m1_subset_1(X3,k2_zfmisc_1(sK0,X1))) & m1_subset_1(X2,k2_zfmisc_1(sK0,X1))) & ~v1_xboole_0(X1)) => (? [X2] : (? [X3] : (X2 != X3 & k2_xtuple_0(X2) = k2_xtuple_0(X3) & k1_xtuple_0(X2) = k1_xtuple_0(X3) & m1_subset_1(X3,k2_zfmisc_1(sK0,sK1))) & m1_subset_1(X2,k2_zfmisc_1(sK0,sK1))) & ~v1_xboole_0(sK1)) [choice axiom]
72. ? [X2] : (? [X3] : (X2 != X3 & k2_xtuple_0(X2) = k2_xtuple_0(X3) & k1_xtuple_0(X2) = k1_xtuple_0(X3) & m1_subset_1(X3,k2_zfmisc_1(sK0,sK1))) & m1_subset_1(X2,k2_zfmisc_1(sK0,sK1))) => (? [X3] : (sK2 != X3 & k2_xtuple_0(X3) = k2_xtuple_0(sK2) & k1_xtuple_0(X3) = k1_xtuple_0(sK2) & m1_subset_1(X3,k2_zfmisc_1(sK0,sK1))) & m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1))) [choice axiom]
73. ? [X3] : (sK2 != X3 & k2_xtuple_0(X3) = k2_xtuple_0(sK2) & k1_xtuple_0(X3) = k1_xtuple_0(sK2) & m1_subset_1(X3,k2_zfmisc_1(sK0,sK1))) => (sK2 != sK3 & k2_xtuple_0(sK2) = k2_xtuple_0(sK3) & k1_xtuple_0(sK2) = k1_xtuple_0(sK3) & m1_subset_1(sK3,k2_zfmisc_1(sK0,sK1))) [choice axiom]
74. (((sK2 != sK3 & k2_xtuple_0(sK2) = k2_xtuple_0(sK3) & k1_xtuple_0(sK2) = k1_xtuple_0(sK3) & m1_subset_1(sK3,k2_zfmisc_1(sK0,sK1))) & m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1))) & ~v1_xboole_0(sK1)) & ~v1_xboole_0(sK0) [skolemisation 52,73,72,71,70]
87. ~v1_xboole_0(sK0) [cnf transformation 74] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
88. ~v1_xboole_0(sK1) [cnf transformation 74] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
89. m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1)) [cnf transformation 74] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
90. m1_subset_1(sK3,k2_zfmisc_1(sK0,sK1)) [cnf transformation 74] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
91. k1_xtuple_0(sK2) = k1_xtuple_0(sK3) [cnf transformation 74] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
92. k2_xtuple_0(sK2) = k2_xtuple_0(sK3) [cnf transformation 74] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
93. sK2 != sK3 [cnf transformation 74] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
98. k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0) [cnf transformation 53] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
104. k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~r2_hidden(X0,X1) | ~v1_relat_1(X1) [cnf transformation 61] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
106. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 64] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
111. v1_xtuple_0(k4_tarski(X0,X1)) [cnf transformation 22] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
112. v1_relat_1(k2_zfmisc_1(X0,X1)) [cnf transformation 26] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
118. ~v1_xboole_0(k2_zfmisc_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 69] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
127. 1 <=> sK2 = sK3 [avatar definition]
129. sK2 != sK3 <- (~1) [avatar component clause 127] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
130. ~1 [avatar split clause 93,127]
132. 2 <=> k2_xtuple_0(sK2) = k2_xtuple_0(sK3) [avatar definition]
134. k2_xtuple_0(sK2) = k2_xtuple_0(sK3) <- (2) [avatar component clause 132] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
135. 2 [avatar split clause 92,132]
137. 3 <=> k1_xtuple_0(sK2) = k1_xtuple_0(sK3) [avatar definition]
139. k1_xtuple_0(sK2) = k1_xtuple_0(sK3) <- (3) [avatar component clause 137] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
140. 3 [avatar split clause 91,137]
142. 4 <=> m1_subset_1(sK3,k2_zfmisc_1(sK0,sK1)) [avatar definition]
144. m1_subset_1(sK3,k2_zfmisc_1(sK0,sK1)) <- (4) [avatar component clause 142] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
145. 4 [avatar split clause 90,142]
147. 5 <=> m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1)) [avatar definition]
149. m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1)) <- (5) [avatar component clause 147] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
150. 5 [avatar split clause 89,147]
152. 6 <=> v1_xboole_0(sK1) [avatar definition]
154. ~v1_xboole_0(sK1) <- (~6) [avatar component clause 152] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
155. ~6 [avatar split clause 88,152]
157. 7 <=> v1_xboole_0(sK0) [avatar definition]
159. ~v1_xboole_0(sK0) <- (~7) [avatar component clause 157] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
160. ~7 [avatar split clause 87,157]
224. v1_xboole_0(k2_zfmisc_1(sK0,sK1)) | r2_hidden(sK3,k2_zfmisc_1(sK0,sK1)) <- (4) [resolution 106,144] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
225. v1_xboole_0(k2_zfmisc_1(sK0,sK1)) | r2_hidden(sK2,k2_zfmisc_1(sK0,sK1)) <- (5) [resolution 106,149] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
228. 19 <=> r2_hidden(sK3,k2_zfmisc_1(sK0,sK1)) [avatar definition]
230. r2_hidden(sK3,k2_zfmisc_1(sK0,sK1)) <- (19) [avatar component clause 228] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
232. 20 <=> v1_xboole_0(k2_zfmisc_1(sK0,sK1)) [avatar definition]
234. v1_xboole_0(k2_zfmisc_1(sK0,sK1)) <- (20) [avatar component clause 232] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
235. 19 | 20 | ~4 [avatar split clause 224,142,232,228]
237. 21 <=> r2_hidden(sK2,k2_zfmisc_1(sK0,sK1)) [avatar definition]
239. r2_hidden(sK2,k2_zfmisc_1(sK0,sK1)) <- (21) [avatar component clause 237] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
240. 21 | 20 | ~5 [avatar split clause 225,147,232,237]
242. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (20) [resolution 234,118] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
245. v1_xboole_0(sK0) <- (~6, 20) [subsumption resolution 242,154] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
246. $false <- (~6, ~7, 20) [subsumption resolution 245,159] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
247. 6 | 7 | ~20 [avatar contradiction clause 246]
264. 23 <=> sK3 = k4_tarski(k1_xtuple_0(sK2),k2_xtuple_0(sK2)) [avatar definition]
266. sK3 = k4_tarski(k1_xtuple_0(sK2),k2_xtuple_0(sK2)) <- (23) [avatar component clause 264] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
282. sK3 = k4_tarski(k1_xtuple_0(sK3),k2_xtuple_0(sK2)) | ~r2_hidden(sK3,X0) | ~v1_relat_1(X0) <- (2) [superposition 104,134] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
287. ~r2_hidden(X13,X14) | v1_xtuple_0(X13) | ~v1_relat_1(X14) [superposition 111,104] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
290. 24 <=> ! [X0] : (~r2_hidden(sK3,X0) | ~v1_relat_1(X0)) [avatar definition]
291. ~r2_hidden(sK3,X0) | ~v1_relat_1(X0) <- (24) [avatar component clause 290] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
293. sK3 = k4_tarski(k1_xtuple_0(sK2),k2_xtuple_0(sK2)) | ~r2_hidden(sK3,X0) | ~v1_relat_1(X0) <- (2, 3) [forward demodulation 282,139] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
294. 24 | 23 | ~2 | ~3 [avatar split clause 293,137,132,264,290]
312. ~v1_relat_1(k2_zfmisc_1(sK0,sK1)) <- (19, 24) [resolution 291,230] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:12,thDist:-12}
314. $false <- (19, 24) [subsumption resolution 312,112] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
315. ~19 | ~24 [avatar contradiction clause 314]
319. sK2 = sK3 | ~v1_xtuple_0(sK2) <- (23) [superposition 98,266] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
329. 27 <=> v1_xtuple_0(sK2) [avatar definition]
331. ~v1_xtuple_0(sK2) <- (~27) [avatar component clause 329] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
334. ~v1_xtuple_0(sK2) <- (~1, 23) [subsumption resolution 319,129] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
335. ~27 | 1 | ~23 [avatar split clause 334,264,127,329]
358. v1_xtuple_0(sK2) | ~v1_relat_1(k2_zfmisc_1(sK0,sK1)) <- (21) [resolution 287,239] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
361. ~v1_relat_1(k2_zfmisc_1(sK0,sK1)) <- (21, ~27) [subsumption resolution 358,331] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
362. $false <- (21, ~27) [subsumption resolution 361,112] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
363. ~21 | 27 [avatar contradiction clause 362]
364. $false [avatar sat refutation 130,135,140,145,150,155,160,235,240,247,294,315,335,363] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:279,thDist:-279}
% SZS output end Proof for domain_1__t2_domain_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 40
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 39
% Generated clauses: 156
% Active clauses: 54
% Passive clauses: 76
% Final active clauses: 50
% Final passive clauses: 21
% Main loop iterations started: 55
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 12
% Fw demodulations: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 16
% Fw demodulations to eq. taut.: 4
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 3
% Forward superposition: 19
% Backward superposition: 23
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 40
% SAT solver unit clauses: 16
% SAT solver binary clauses: 12
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% SAT solver time: 0.001 s
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

