i: [1,56,0,2,1,1,0,0] t12_funct_9
a: [3,95,0,2,1,56]
s: 95
i: [1,55,0,2,1,1,0,0] t12_funct_9
a: [3,100,0,2,1,55]
s: 100
i: [1,54,0,2,1,1,0,0] t12_funct_9
a: [3,105,0,2,1,54]
s: 105
i: [1,53,0,2,1,1,0,0] t12_funct_9
a: [3,110,0,2,1,53]
s: 110
i: [1,57,0,3,1,1,0,0] t12_funct_9
a: [3,90,0,3,1,57]
s: 90
i: [1,58,0,4,1,1,0,0] t12_funct_9
a: [3,85,0,4,1,58]
s: 85
i: [1,81,0,2,1,0,0,1] rc1_xreal_0
a: [3,115,0,2,1,81]
s: 115
i: [1,62,0,6,3,0,0,1] cc6_valued_0
s: 62
d: [2,117,1,4,2,1,40,62,95]
d: [2,118,1,2,1,2,21,117,105]
a: [3,122,1,2,1,118]
s: 122
i: [1,63,0,9,4,0,0,1] fc71_valued_1
s: 63
i: [1,64,0,9,4,0,0,1] fc71_valued_1
s: 64
i: [1,65,0,9,4,0,0,1] fc71_valued_1
s: 65
i: [1,67,0,9,4,0,0,1] dt_k39_valued_1
s: 67
i: [1,68,0,9,4,0,0,1] dt_k39_valued_1
s: 68
i: [1,69,0,9,4,0,0,1] fc70_valued_1
s: 69
i: [1,70,0,9,4,0,0,1] fc70_valued_1
s: 70
i: [1,71,0,9,4,0,0,1] fc70_valued_1
s: 71
i: [1,66,0,12,4,0,0,255] d8_valued_1
s: 66
d: [2,124,1,12,5,0,40,65,62]
d: [2,125,1,9,4,0,21,124,63]
s: 125
i: [1,72,0,16,7,0,0,1] fc34_valued_1
s: 72
i: [1,73,0,16,7,0,0,1] fc34_valued_1
s: 73
i: [1,74,0,16,7,0,0,1] fc34_valued_1
s: 74
i: [1,76,0,16,7,0,0,2] dt_k18_valued_1
s: 76
i: [1,77,0,16,7,0,0,2] dt_k18_valued_1
s: 77
i: [1,78,0,16,7,0,0,2] fc33_valued_1
s: 78
i: [1,79,0,16,7,0,0,2] fc33_valued_1
s: 79
i: [1,80,0,16,7,0,0,2] fc33_valued_1
s: 80
i: [1,75,0,19,7,0,0,255] commutativity_k18_valued_1
s: 75
d: [2,130,1,20,8,0,40,74,62]
d: [2,133,1,16,7,0,21,130,72]
s: 133
i: [1,59,0,25,10,0,0,1] t4_funct_9
s: 59
d: [2,170,1,30,13,0,44,59,66]
d: [2,175,1,17,7,0,18,170]
d: [2,176,1,15,6,0,21,175,62]
s: 176
d: [2,177,2,0,0,6,52,110,105,100,95,90,85,176]
e: 177
d: [2,178,2,11,5,1,40,176,85]
d: [2,180,2,8,4,2,21,178,90]
d: [2,181,2,6,3,3,21,180,95]
d: [2,182,2,4,2,4,21,181,100]
d: [2,183,2,2,1,5,21,182,105]
d: [2,184,2,0,0,6,21,183,110]
e: 184
f: 184,53,54,55,56,57,58
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_9__t12_funct_9
% SZS output start Proof for funct_9__t12_funct_9
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : ((v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => (v1_funct_9(X1,X0) => v1_funct_9(k39_valued_1(X1),X0)))) [input t12_funct_9]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : ((v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => (v1_funct_9(X1,X0) => v1_funct_9(k39_valued_1(X1),X0)))) [negated conjecture 1]
13. ! [X0] : ((v3_valued_0(X0) & v1_relat_1(X0)) => (v1_valued_0(X0) & v1_relat_1(X0))) [input cc6_valued_0]
17. ! [X0] : ((v1_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => k39_valued_1(X0) = k18_valued_1(X0,X0)) [input d8_valued_1]
18. ! [X0] : (v1_xreal_0(X0) => ! [X1] : ((v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((v1_funct_9(X2,X0) & v1_funct_9(X1,X0)) => v1_funct_9(k18_valued_1(X1,X2),X0))))) [input t4_funct_9]
24. ? [X0] : (? [X1] : ((~v1_funct_9(k39_valued_1(X1),X0) & v1_funct_9(X1,X0)) & (v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1))) & v1_xreal_0(X0)) [ennf transformation 2]
25. ? [X0] : (? [X1] : (~v1_funct_9(k39_valued_1(X1),X0) & v1_funct_9(X1,X0) & v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_xreal_0(X0)) [flattening 24]
26. ! [X0] : (! [X1] : (! [X2] : ((v1_funct_9(k18_valued_1(X1,X2),X0) | (~v1_funct_9(X2,X0) | ~v1_funct_9(X1,X0))) | (~v3_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v3_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1))) | ~v1_xreal_0(X0)) [ennf transformation 18]
27. ! [X0] : (! [X1] : (! [X2] : (v1_funct_9(k18_valued_1(X1,X2),X0) | ~v1_funct_9(X2,X0) | ~v1_funct_9(X1,X0) | ~v3_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v3_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_xreal_0(X0)) [flattening 26]
30. ! [X0] : ((v1_valued_0(X0) & v1_relat_1(X0)) | (~v3_valued_0(X0) | ~v1_relat_1(X0))) [ennf transformation 13]
31. ! [X0] : ((v1_valued_0(X0) & v1_relat_1(X0)) | ~v3_valued_0(X0) | ~v1_relat_1(X0)) [flattening 30]
34. ! [X0] : (k39_valued_1(X0) = k18_valued_1(X0,X0) | (~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 17]
35. ! [X0] : (k39_valued_1(X0) = k18_valued_1(X0,X0) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 34]
48. ? [X0] : (? [X1] : (~v1_funct_9(k39_valued_1(X1),X0) & v1_funct_9(X1,X0) & v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_xreal_0(X0)) => (? [X1] : (~v1_funct_9(k39_valued_1(X1),sK0) & v1_funct_9(X1,sK0) & v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_xreal_0(sK0)) [choice axiom]
49. ? [X1] : (~v1_funct_9(k39_valued_1(X1),sK0) & v1_funct_9(X1,sK0) & v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => (~v1_funct_9(k39_valued_1(sK1),sK0) & v1_funct_9(sK1,sK0) & v3_valued_0(sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
50. (~v1_funct_9(k39_valued_1(sK1),sK0) & v1_funct_9(sK1,sK0) & v3_valued_0(sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_xreal_0(sK0) [skolemisation 25,49,48]
53. v1_xreal_0(sK0) [cnf transformation 50] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. v1_relat_1(sK1) [cnf transformation 50] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. v1_funct_1(sK1) [cnf transformation 50] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. v3_valued_0(sK1) [cnf transformation 50] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
57. v1_funct_9(sK1,sK0) [cnf transformation 50] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. ~v1_funct_9(k39_valued_1(sK1),sK0) [cnf transformation 50] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
59. v1_funct_9(k18_valued_1(X1,X2),X0) | ~v1_funct_9(X2,X0) | ~v1_funct_9(X1,X0) | ~v3_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v3_valued_0(X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_xreal_0(X0) [cnf transformation 27] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. ~v3_valued_0(X0) | v1_valued_0(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
66. k39_valued_1(X0) = k18_valued_1(X0,X0) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
83. 1 <=> v1_funct_9(k39_valued_1(sK1),sK0) [avatar definition]
85. ~v1_funct_9(k39_valued_1(sK1),sK0) <- (~1) [avatar component clause 83] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. ~1 [avatar split clause 58,83]
88. 2 <=> v1_funct_9(sK1,sK0) [avatar definition]
90. v1_funct_9(sK1,sK0) <- (2) [avatar component clause 88] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. 2 [avatar split clause 57,88]
93. 3 <=> v3_valued_0(sK1) [avatar definition]
95. v3_valued_0(sK1) <- (3) [avatar component clause 93] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. 3 [avatar split clause 56,93]
98. 4 <=> v1_funct_1(sK1) [avatar definition]
100. v1_funct_1(sK1) <- (4) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. 4 [avatar split clause 55,98]
103. 5 <=> v1_relat_1(sK1) [avatar definition]
105. v1_relat_1(sK1) <- (5) [avatar component clause 103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
106. 5 [avatar split clause 54,103]
108. 6 <=> v1_xreal_0(sK0) [avatar definition]
110. v1_xreal_0(sK0) <- (6) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
111. 6 [avatar split clause 53,108]
170. v1_funct_9(k39_valued_1(X0),X1) | ~v1_funct_9(X0,X1) | ~v1_funct_9(X0,X1) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_xreal_0(X1) | ~v1_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [superposition 59,66] {a:1,w:30,wCS:12000,sine:1,thAx:0,allAx:2,thDist:-2}
175. v1_funct_9(k39_valued_1(X0),X1) | ~v1_funct_9(X0,X1) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_xreal_0(X1) | ~v1_valued_0(X0) [duplicate literal removal 170] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
176. v1_funct_9(k39_valued_1(X0),X1) | ~v1_funct_9(X0,X1) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_xreal_0(X1) [subsumption resolution 175,62] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
178. ~v1_funct_9(sK1,sK0) | ~v3_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_xreal_0(sK0) <- (~1) [resolution 176,85] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
180. ~v3_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_xreal_0(sK0) <- (~1, 2) [subsumption resolution 178,90] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
181. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_xreal_0(sK0) <- (~1, 2, 3) [subsumption resolution 180,95] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
182. ~v1_relat_1(sK1) | ~v1_xreal_0(sK0) <- (~1, 2, 3, 4) [subsumption resolution 181,100] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
183. ~v1_xreal_0(sK0) <- (~1, 2, 3, 4, 5) [subsumption resolution 182,105] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
184. $false <- (~1, 2, 3, 4, 5, 6) [subsumption resolution 183,110] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
185. 1 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar contradiction clause 184]
186. $false [avatar sat refutation 86,91,96,101,106,111,185] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:59,thDist:-59}
% SZS output end Proof for funct_9__t12_funct_9
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% >>> Preprocessing
% Introduced skolems: 3
% Pure predicates: 3
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 99
% Active clauses: 27
% Passive clauses: 31
% Final active clauses: 27
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Duplicate literals: 120
% Fw subsumption resolutions: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 27
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% Forward superposition: 12
% Backward superposition: 12
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 7
% 
% Memory used [KB]: 9338
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

