i: [1,51,0,14,3,1,0,0] t106_zfmisc_1
a: [3,82,0,3,1,51]
s: 82
a: [3,86,0,3,1,51]
s: 86
a: [3,79,0,8,1,51]
s: 79
d: [2,116,0,8,1,2,22,79,86]
a: [3,120,0,8,1,116]
s: 120
i: [1,52,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,94,0,2,1,52]
s: 94
i: [1,53,0,2,1,0,0,255] fc1_xboole_0
a: [3,99,0,2,1,53]
s: 99
i: [1,71,0,2,1,0,0,2] rc2_xboole_0
a: [3,109,0,2,1,71]
s: 109
i: [1,72,0,2,1,0,0,2] rc1_xboole_0
a: [3,114,0,2,1,72]
s: 114
i: [1,55,0,3,1,0,0,1] fc2_xboole_0
s: 55
i: [1,54,0,3,1,0,0,255] d2_xboole_0
a: [3,104,0,3,1,54]
s: 104
r: 94
i: [1,67,0,4,1,0,0,1] fc1_zfmisc_1
s: 67
i: [1,68,0,4,1,0,0,255] fc3_xboole_0
s: 68
i: [1,74,0,4,1,0,0,1] d1_tarski
s: 74
i: [1,56,0,5,2,0,0,255] t6_boole
s: 56
i: [1,63,0,5,2,0,0,1] t7_boole
s: 63
i: [1,70,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 70
d: [2,125,1,2,1,1,40,63,82]
a: [3,130,1,2,1,125]
s: 130
i: [1,62,0,7,3,0,0,1] t8_boole
s: 62
d: [2,132,1,3,1,1,40,70,82]
a: [3,137,1,3,1,132]
s: 137
i: [1,69,0,7,1,0,0,255] commutativity_k2_tarski
s: 69
i: [1,75,0,7,2,0,0,1] d1_tarski
s: 75
d: [2,133,1,4,1,0,40,70,74]
s: 133
i: [1,57,0,10,1,0,0,255] d5_tarski
s: 57
d: [2,123,1,7,3,0,44,56,56]
s: 123
i: [1,64,0,10,2,0,0,1] l46_zfmisc_1
s: 64
i: [1,65,0,10,2,0,0,1] l46_zfmisc_1
s: 65
d: [2,142,1,10,1,0,44,57,69]
s: 142
i: [1,66,0,13,3,0,0,1] l46_zfmisc_1
s: 66
d: [2,160,1,0,0,2,52,82,74,120,66]
e: 160
d: [2,161,1,7,2,1,40,66,120]
d: [2,167,1,3,1,1,21,161,74]
d: [2,168,1,0,0,2,21,167,82]
e: 168
r: 86
r: 120
a: [3,87,0,3,1,51]
s: 87
a: [3,78,0,8,1,51]
s: 78
i: [1,60,0,14,3,0,0,1] d1_tarski
s: 60
d: [2,170,1,4,1,1,40,78,65]
a: [3,177,1,4,1,170]
s: 177
d: [2,192,2,0,0,2,52,87,177,75]
e: 192
d: [2,193,2,3,1,1,40,177,75]
d: [2,197,2,0,0,2,21,193,87]
e: 197
r: 82
r: 87
i: [1,50,0,11,2,1,0,0] t106_zfmisc_1
i: [1,49,0,11,2,1,0,0] t106_zfmisc_1
d: [2,171,1,3,1,1,40,78,64]
f: 193,171,170,168,116,49,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t106_zfmisc_1
% SZS output start Proof for zfmisc_1__t106_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3))) <=> (X1 = X3 & r2_hidden(X0,X2))) [input t106_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3))) <=> (X1 = X3 & r2_hidden(X0,X2))) [negated conjecture 1]
5. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
18. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [input l46_zfmisc_1]
24. ~! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3))) <=> (X1 = X3 & r2_hidden(X0,X2))) [flattening 2]
26. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 5]
29. ! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [flattening 18]
30. ? [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3))) <~> (X1 = X3 & r2_hidden(X0,X2))) [ennf transformation 24]
35. ? [X0,X1,X2,X3] : (((X1 != X3 | ~r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3)))) & ((X1 = X3 & r2_hidden(X0,X2)) | r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3))))) [nnf transformation 30]
36. ? [X0,X1,X2,X3] : ((X1 != X3 | ~r2_hidden(X0,X2) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3)))) & ((X1 = X3 & r2_hidden(X0,X2)) | r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3))))) [flattening 35]
37. ? [X0,X1,X2,X3] : ((X1 != X3 | ~r2_hidden(X0,X2) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3)))) & ((X1 = X3 & r2_hidden(X0,X2)) | r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,k1_tarski(X3))))) => ((sK1 != sK3 | ~r2_hidden(sK0,sK2) | ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3)))) & ((sK1 = sK3 & r2_hidden(sK0,sK2)) | r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3))))) [choice axiom]
38. (sK1 != sK3 | ~r2_hidden(sK0,sK2) | ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3)))) & ((sK1 = sK3 & r2_hidden(sK0,sK2)) | r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3)))) [skolemisation 36,37]
39. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 26]
40. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 39]
41. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 40,41]
43. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | (~r2_hidden(X1,X3) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [nnf transformation 29]
44. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [flattening 43]
49. r2_hidden(sK0,sK2) | r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3))) [cnf transformation 38] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. sK1 = sK3 | r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3))) [cnf transformation 38] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. sK1 != sK3 | ~r2_hidden(sK0,sK2) | ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3))) [cnf transformation 38] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 42] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
59. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 42] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
64. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X0,X2) [cnf transformation 44] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X1,X3) [cnf transformation 44] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
66. r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2) [cnf transformation 44] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 59] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:2,thDist:-2}
74. r2_hidden(X3,k1_tarski(X3)) [equality resolution 73] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
75. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 58] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
77. 1 <=> r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3))) [avatar definition]
78. r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3))) <- (1) [avatar component clause 77] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK3))) <- (~1) [avatar component clause 77] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
81. 2 <=> r2_hidden(sK0,sK2) [avatar definition]
82. r2_hidden(sK0,sK2) <- (2) [avatar component clause 81] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
85. 3 <=> sK1 = sK3 [avatar definition]
86. sK1 = sK3 <- (3) [avatar component clause 85] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
88. ~1 | ~2 | ~3 [avatar split clause 51,85,81,77]
89. 1 | 3 [avatar split clause 50,85,77]
90. 1 | 2 [avatar split clause 49,81,77]
116. ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK1))) <- (~1, 3) [forward demodulation 79,86] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
118. 9 <=> r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK1))) [avatar definition]
120. ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(sK2,k1_tarski(sK1))) <- (~9) [avatar component clause 118] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
121. ~9 | 1 | ~3 [avatar split clause 116,85,77,118]
161. ~r2_hidden(sK1,k1_tarski(sK1)) | ~r2_hidden(sK0,sK2) <- (~9) [resolution 66,120] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
167. ~r2_hidden(sK0,sK2) <- (~9) [subsumption resolution 161,74] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
168. $false <- (2, ~9) [subsumption resolution 167,82] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
169. ~2 | 9 [avatar contradiction clause 168]
170. r2_hidden(sK1,k1_tarski(sK3)) <- (1) [resolution 78,65] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
171. r2_hidden(sK0,sK2) <- (1) [resolution 78,64] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
175. 12 <=> r2_hidden(sK1,k1_tarski(sK3)) [avatar definition]
177. r2_hidden(sK1,k1_tarski(sK3)) <- (12) [avatar component clause 175] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
178. 12 | ~1 [avatar split clause 170,77,175]
193. sK1 = sK3 <- (12) [resolution 177,75] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
199. 2 | ~1 [avatar split clause 171,77,81]
200. 3 | ~12 [avatar split clause 193,175,85]
206. $false [avatar sat refutation 88,89,90,121,169,178,199,200] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:95,thDist:-95}
% SZS output end Proof for zfmisc_1__t106_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 96
% Active clauses: 30
% Passive clauses: 55
% Final active clauses: 25
% Final passive clauses: 23
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 3
% Fw demodulations: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 2
% Forward superposition: 9
% Backward superposition: 10
% Self superposition: 3
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 5
% SAT solver binary clauses: 14
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

