i: [1,200,0,2,1,1,0,0] t67_xxreal_1
a: [3,320,0,2,1,200]
s: 320
i: [1,199,0,2,1,1,0,0] t67_xxreal_1
a: [3,325,0,2,1,199]
s: 325
i: [1,198,0,2,1,1,0,0] t67_xxreal_1
a: [3,330,0,2,1,198]
s: 330
i: [1,197,0,2,1,1,0,0] t67_xxreal_1
a: [3,335,0,2,1,197]
s: 335
i: [1,203,0,6,2,1,0,0] t67_xxreal_1
a: [3,301,0,3,1,203]
s: 301
i: [1,202,0,3,1,1,0,0] t67_xxreal_1
a: [3,310,0,3,1,202]
s: 310
i: [1,201,0,7,1,1,0,0] t67_xxreal_1
a: [3,315,0,7,1,201]
s: 315
i: [1,204,0,2,1,0,0,255] fc5_numbers
a: [3,340,0,2,1,204]
s: 340
i: [1,205,0,2,1,0,0,255] fc2_membered
a: [3,345,0,2,1,205]
s: 345
i: [1,206,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,350,0,2,1,206]
s: 350
i: [1,207,0,2,1,0,0,255] fc1_xboole_0
a: [3,355,0,2,1,207]
s: 355
i: [1,264,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,364,0,2,1,264]
s: 364
d: [2,498,1,0,0,2,40,364,320]
e: 498
d: [2,497,1,0,0,2,40,364,325]
e: 497
d: [2,496,1,0,0,2,40,364,330]
e: 496
d: [2,495,1,0,0,2,40,364,335]
e: 495
d: [2,494,1,0,0,2,52,335,364]
e: 494
r: 364
i: [1,272,0,2,1,0,0,2] rc2_xboole_0
a: [3,372,0,2,1,272]
s: 372
i: [1,273,0,2,1,0,0,255] rc1_xreal_0
a: [3,377,0,2,1,273]
s: 377
i: [1,274,0,2,1,0,0,2] rc1_xboole_0
a: [3,382,0,2,1,274]
s: 382
i: [1,275,0,2,1,0,0,1] rc1_xxreal_0
a: [3,387,0,2,1,275]
s: 387
i: [1,277,0,2,1,0,0,255] rc1_membered
a: [3,392,0,2,1,277]
s: 392
i: [1,276,0,2,1,0,0,255] rc1_membered
a: [3,397,0,2,1,276]
s: 397
i: [1,279,0,2,1,0,0,255] rc2_membered
a: [3,402,0,2,1,279]
s: 402
i: [1,278,0,2,1,0,0,255] rc2_membered
a: [3,407,0,2,1,278]
s: 407
i: [1,281,0,2,1,0,0,255] rc3_membered
a: [3,412,0,2,1,281]
s: 412
i: [1,280,0,2,1,0,0,255] rc3_membered
a: [3,417,0,2,1,280]
s: 417
i: [1,284,0,2,1,0,0,255] rc3_xreal_0
a: [3,422,0,2,1,284]
s: 422
i: [1,283,0,2,1,0,0,255] rc3_xreal_0
a: [3,427,0,2,1,283]
s: 427
i: [1,282,0,2,1,0,0,255] rc3_xreal_0
a: [3,432,0,2,1,282]
s: 432
i: [1,287,0,2,1,0,0,255] rc2_xreal_0
a: [3,437,0,2,1,287]
s: 437
i: [1,286,0,2,1,0,0,255] rc2_xreal_0
a: [3,442,0,2,1,286]
s: 442
i: [1,285,0,2,1,0,0,255] rc2_xreal_0
a: [3,447,0,2,1,285]
s: 447
i: [1,290,0,2,1,0,0,255] rc4_xreal_0
a: [3,452,0,2,1,290]
s: 452
i: [1,289,0,2,1,0,0,255] rc4_xreal_0
a: [3,457,0,2,1,289]
s: 457
i: [1,288,0,2,1,0,0,255] rc4_xreal_0
a: [3,462,0,2,1,288]
s: 462
i: [1,292,0,2,1,0,0,2] rc4_xxreal_0
a: [3,467,0,2,1,292]
s: 467
i: [1,291,0,2,1,0,0,2] rc4_xxreal_0
a: [3,472,0,2,1,291]
s: 472
i: [1,294,0,2,1,0,0,255] rc2_xxreal_0
a: [3,477,0,2,1,294]
s: 477
i: [1,293,0,2,1,0,0,255] rc2_xxreal_0
a: [3,482,0,2,1,293]
s: 482
i: [1,296,0,2,1,0,0,255] rc3_xxreal_0
a: [3,487,0,2,1,296]
s: 487
i: [1,295,0,2,1,0,0,255] rc3_xxreal_0
a: [3,492,0,2,1,295]
s: 492
i: [1,262,0,3,1,0,0,255] reflexivity_r1_tarski
s: 262
i: [1,208,0,3,1,0,0,255] d2_xboole_0
a: [3,360,0,3,1,208]
s: 360
r: 350
i: [1,209,0,4,2,0,0,255] cc2_xreal_0
s: 209
i: [1,210,0,4,2,0,0,255] cc2_xxreal_0
s: 210
i: [1,212,0,4,2,0,0,255] cc2_membered
s: 212
i: [1,213,0,4,2,0,0,255] cc3_membered
s: 213
i: [1,216,0,4,2,0,0,255] cc5_membered
s: 216
i: [1,217,0,4,2,0,0,255] cc4_membered
s: 217
i: [1,222,0,4,2,0,0,255] cc1_membered
s: 222
d: [2,505,1,2,1,1,40,222,392]
a: [3,511,1,2,1,505]
s: 511
i: [1,225,0,4,2,0,0,255] cc4_xreal_0
s: 225
d: [2,506,1,2,1,1,40,222,402]
a: [3,516,1,2,1,506]
s: 516
i: [1,226,0,4,2,0,0,255] cc18_membered
s: 226
i: [1,261,0,4,1,0,0,3] existence_m1_subset_1
s: 261
d: [2,507,1,2,1,1,40,222,412]
a: [3,521,1,2,1,507]
s: 521
i: [1,227,0,5,2,0,0,255] t6_boole
s: 227
d: [2,529,1,2,1,1,40,225,377]
a: [3,536,1,2,1,529]
s: 536
i: [1,234,0,5,2,0,0,255] cc1_xxreal_0
s: 234
i: [1,256,0,5,2,0,0,255] t7_boole
s: 256
d: [2,544,1,2,1,1,40,226,355]
a: [3,551,1,2,1,544]
s: 551
a: [3,367,0,5,2,264]
s: 367
d: [2,545,1,2,1,1,40,226,382]
a: [3,556,1,2,1,545]
s: 556
i: [1,228,0,6,2,0,0,1] fc8_xxreal_1
s: 228
i: [1,235,0,6,2,0,0,255] cc7_membered
s: 235
d: [2,546,1,2,1,1,40,226,462]
a: [3,561,1,2,1,546]
s: 561
i: [1,241,0,6,3,0,0,255] cc7_xxreal_0
s: 241
d: [2,547,1,2,1,1,40,226,472]
a: [3,566,1,2,1,547]
s: 566
i: [1,242,0,6,3,0,0,255] cc7_xxreal_0
s: 242
i: [1,245,0,6,3,0,0,255] cc3_xxreal_0
s: 245
d: [2,608,1,4,2,1,40,241,442]
d: [2,610,1,2,1,2,21,608,447]
a: [3,614,1,2,1,610]
s: 614
i: [1,247,0,6,3,0,0,255] cc3_xxreal_0
s: 247
d: [2,609,1,4,2,1,40,241,477]
d: [2,616,1,2,1,2,21,609,482]
a: [3,620,1,2,1,616]
s: 620
i: [1,248,0,6,3,0,0,255] cc5_xxreal_0
s: 248
i: [1,250,0,6,3,0,0,255] cc5_xxreal_0
s: 250
d: [2,628,1,4,2,1,40,242,427]
d: [2,630,1,2,1,2,21,628,432]
a: [3,634,1,2,1,630]
s: 634
i: [1,251,0,6,2,0,0,255] t1_subset
s: 251
d: [2,629,1,4,2,1,40,242,487]
d: [2,636,1,2,1,2,21,629,492]
a: [3,640,1,2,1,636]
s: 640
i: [1,263,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 263
i: [1,218,0,7,3,0,0,255] cc14_membered
s: 218
d: [2,642,1,4,2,1,40,247,427]
d: [2,644,1,2,1,2,21,642,432]
a: [3,648,1,2,1,644]
s: 648
i: [1,220,0,7,3,0,0,2] cc13_membered
s: 220
d: [2,643,1,4,2,1,40,247,487]
d: [2,650,1,2,1,2,21,643,492]
a: [3,654,1,2,1,650]
s: 654
i: [1,223,0,7,3,0,0,255] cc17_membered
s: 223
i: [1,253,0,7,2,0,0,255] t3_subset
s: 253
d: [2,523,2,2,1,1,40,511,212]
a: [3,527,2,2,1,523]
s: 527
i: [1,254,0,7,2,0,0,255] t3_subset
s: 254
d: [2,538,2,2,1,1,40,516,212]
a: [3,542,2,2,1,538]
s: 542
i: [1,255,0,7,3,0,0,1] t8_boole
s: 255
i: [1,257,0,7,2,0,0,255] t2_tarski
s: 257
d: [2,568,2,2,1,1,40,521,212]
a: [3,572,2,2,1,568]
s: 572
i: [1,211,0,8,3,0,0,255] cc19_membered
s: 211
d: [2,582,2,2,1,1,40,551,222]
a: [3,587,2,2,1,582]
s: 587
i: [1,214,0,8,3,0,0,255] cc22_membered
s: 214
i: [1,215,0,8,3,0,0,255] cc23_membered
s: 215
d: [2,589,2,2,1,1,40,556,222]
a: [3,593,2,2,1,589]
s: 593
i: [1,219,0,8,3,0,0,255] cc21_membered
s: 219
d: [2,602,2,2,1,1,40,561,222]
a: [3,606,2,2,1,602]
s: 606
i: [1,221,0,8,3,0,0,2] cc20_membered
s: 221
i: [1,224,0,8,3,0,0,255] cc24_membered
s: 224
d: [2,622,2,2,1,1,40,566,222]
a: [3,626,2,2,1,622]
s: 626
i: [1,237,0,8,4,0,0,255] cc4_xxreal_0
s: 237
d: [2,723,1,6,3,1,40,237,335]
a: [3,738,1,2,1,723]
s: 738
r: 738
i: [1,239,0,8,4,0,0,255] cc6_xxreal_0
s: 239
i: [1,243,0,8,4,0,0,255] cc8_xxreal_0
s: 243
a: [3,737,1,2,1,723]
s: 737
i: [1,252,0,8,3,0,0,255] t2_subset
s: 252
a: [3,742,1,2,1,723]
s: 742
i: [1,265,0,8,3,0,0,1] fc4_xxreal_1
s: 265
i: [1,260,0,9,3,0,0,255] t5_subset
s: 260
d: [2,724,1,6,3,1,40,237,330]
a: [3,751,1,2,1,724]
s: 751
r: 751
i: [1,258,0,10,3,0,0,255] t2_tarski
s: 258
a: [3,750,1,2,1,724]
s: 750
i: [1,259,0,10,3,0,0,255] t4_subset
s: 259
i: [1,266,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 266
a: [3,755,1,2,1,724]
s: 755
i: [1,229,0,11,3,0,0,255] d4_xxreal_1
s: 229
d: [2,725,1,6,3,1,40,237,325]
a: [3,764,1,2,1,725]
s: 764
r: 764
i: [1,230,0,13,5,0,0,1] t1_xxreal_0
s: 230
i: [1,231,0,15,6,0,0,1] t2_xxreal_0
s: 231
a: [3,763,1,2,1,725]
s: 763
i: [1,267,0,15,4,0,0,255] fraenkel_a_2_3_xxreal_1
s: 267
a: [3,768,1,2,1,725]
s: 768
i: [1,268,0,15,4,0,0,255] fraenkel_a_2_3_xxreal_1
s: 268
i: [1,269,0,15,4,0,0,255] fraenkel_a_2_3_xxreal_1
s: 269
d: [2,726,1,6,3,1,40,237,320]
a: [3,777,1,2,1,726]
s: 777
r: 777
i: [1,270,0,15,4,0,0,255] fraenkel_a_2_3_xxreal_1
s: 270
a: [3,776,1,2,1,726]
s: 776
i: [1,297,0,18,6,0,0,255] fraenkel_a_2_3_xxreal_1
s: 297
i: [1,232,0,21,7,0,0,255] t63_xxreal_1
s: 232
a: [3,781,1,2,1,726]
s: 781
i: [1,233,0,21,7,0,0,255] t63_xxreal_1
s: 233
d: [2,841,1,7,3,1,40,266,310]
d: [2,844,1,5,2,2,21,841,330]
d: [2,845,1,3,1,3,21,844,335]
a: [3,849,1,3,1,845]
s: 849
d: [2,912,1,21,7,1,44,232,315]
d: [2,915,1,19,6,2,21,912,320]
d: [2,916,1,17,5,3,21,915,325]
a: [3,920,1,3,1,916]
s: 920
d: [2,822,1,8,3,1,44,265,315]
d: [2,823,1,6,2,2,21,822,320]
d: [2,824,1,4,1,3,21,823,325]
a: [3,828,1,4,1,824]
s: 828
d: [2,945,2,10,4,1,40,920,230]
d: [2,948,2,8,3,2,21,945,320]
d: [2,949,2,6,2,3,21,948,325]
a: [3,953,2,3,1,949]
s: 953
d: [2,728,1,6,3,1,40,237,387]
a: [3,803,1,2,1,728]
s: 803
r: 803
a: [3,802,1,2,1,728]
s: 802
a: [3,807,1,2,1,728]
s: 807
d: [2,868,1,16,5,0,40,267,223]
a: [3,875,1,2,1,868]
s: 875
d: [2,870,1,16,5,0,40,267,218]
a: [3,884,1,2,1,870]
s: 884
d: [2,576,1,3,1,0,40,234,261]
a: [3,580,1,3,1,576]
s: 580
d: [2,940,2,12,5,1,40,849,231]
d: [2,942,2,10,4,2,21,940,330]
d: [2,943,2,8,3,3,21,942,335]
s: 943
d: [2,596,1,4,1,0,40,235,261]
a: [3,600,1,4,1,596]
s: 600
d: [2,913,1,21,7,1,44,232,315]
d: [2,925,1,19,6,2,21,913,320]
d: [2,926,1,17,5,3,21,925,325]
s: 926
d: [2,986,2,10,4,3,40,926,262]
d: [2,988,2,7,3,4,21,986,310]
d: [2,989,2,5,2,5,21,988,330]
d: [2,990,2,3,1,6,21,989,335]
a: [3,994,2,3,1,990]
s: 994
d: [2,930,1,21,7,1,44,233,315]
d: [2,938,1,19,6,2,21,930,320]
d: [2,939,1,17,5,3,21,938,325]
s: 939
d: [2,1015,2,10,4,3,40,939,262]
d: [2,1017,2,7,3,4,21,1015,310]
d: [2,1018,2,5,2,5,21,1017,330]
d: [2,1019,2,3,1,6,21,1018,335]
a: [3,1023,2,3,1,1019]
s: 1023
d: [2,656,1,5,2,0,40,218,261]
s: 656
d: [2,657,1,5,2,0,40,220,261]
s: 657
d: [2,996,3,5,2,4,40,994,943]
d: [2,999,3,3,1,5,21,996,325]
a: [3,1003,3,3,1,999]
s: 1003
d: [2,1038,4,0,0,6,52,325,330,320,1023,953,1003,231]
e: 1038
r: 920
r: 953
a: [3,919,1,3,1,916]
s: 919
a: [3,952,2,3,1,949]
s: 952
a: [3,923,1,14,4,916]
s: 923
r: 301
a: [3,300,0,3,1,203]
s: 300
r: 315
r: 325
r: 763
r: 768
r: 919
r: 923
r: 926
r: 952
r: 994
r: 1003
a: [3,305,0,3,1,203]
s: 305
d: [2,1072,0,7,1,2,23,315,300]
a: [3,1085,0,7,1,1072]
s: 1085
d: [2,1076,1,3,1,2,23,919,300]
a: [3,1090,1,3,1,1076]
s: 1090
d: [2,929,1,21,7,1,44,233,315]
d: [2,932,1,19,6,2,21,929,320]
d: [2,933,1,17,5,3,21,932,325]
a: [3,936,1,14,4,933]
s: 936
d: [2,1026,3,10,4,1,40,1023,230]
d: [2,1029,3,8,3,2,21,1026,330]
d: [2,1030,3,6,2,3,21,1029,320]
d: [2,1063,2,7,3,1,40,923,262]
d: [2,1065,2,5,2,2,21,1063,335]
d: [2,1066,2,3,1,3,21,1065,330]
d: [2,998,3,10,4,1,40,994,230]
d: [2,1069,3,8,3,2,21,998,325]
d: [2,1070,3,6,2,3,21,1069,335]
d: [2,1130,2,7,3,1,40,936,262]
d: [2,1132,2,5,2,2,21,1130,335]
d: [2,1133,2,3,1,3,21,1132,330]
f: 1133,1070,1066,1038,1030,1019,999,990,949,933,916,845,197,198,199,200,201,202,203
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_1__t67_xxreal_1
% SZS output start Proof for xxreal_1__t67_xxreal_1
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ! [X3] : (v1_xxreal_0(X3) => (k4_xxreal_1(X0,X1) = k4_xxreal_1(X2,X3) => ((X1 = X3 & X0 = X2) | r1_xxreal_0(X1,X0))))))) [input t67_xxreal_1]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ! [X3] : (v1_xxreal_0(X3) => (k4_xxreal_1(X0,X1) = k4_xxreal_1(X2,X3) => ((X1 = X3 & X0 = X2) | r1_xxreal_0(X1,X0))))))) [negated conjecture 1]
35. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
64. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
67. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input t1_xxreal_0]
70. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input t2_xxreal_0]
74. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ! [X3] : (v1_xxreal_0(X3) => (r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(X2,X3)) => ((r1_xxreal_0(X1,X3) & r1_xxreal_0(X2,X0)) | r1_xxreal_0(X1,X0))))))) [input t63_xxreal_1]
86. ! [X0] : r1_tarski(X0,X0) [rectify 64]
96. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((X1 != X3 | X0 != X2) & ~r1_xxreal_0(X1,X0)) & k4_xxreal_1(X0,X1) = k4_xxreal_1(X2,X3)) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
97. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((X1 != X3 | X0 != X2) & ~r1_xxreal_0(X1,X0) & k4_xxreal_1(X0,X1) = k4_xxreal_1(X2,X3) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 96]
119. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 67]
120. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 119]
121. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 70]
122. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 121]
123. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((((r1_xxreal_0(X1,X3) & r1_xxreal_0(X2,X0)) | r1_xxreal_0(X1,X0)) | ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(X2,X3))) | ~v1_xxreal_0(X3)) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 74]
124. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_xxreal_0(X1,X3) & r1_xxreal_0(X2,X0)) | r1_xxreal_0(X1,X0) | ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(X2,X3)) | ~v1_xxreal_0(X3)) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 123]
153. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 35]
154. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 153]
157. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((X1 != X3 | X0 != X2) & ~r1_xxreal_0(X1,X0) & k4_xxreal_1(X0,X1) = k4_xxreal_1(X2,X3) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (? [X3] : ((X1 != X3 | sK0 != X2) & ~r1_xxreal_0(X1,sK0) & k4_xxreal_1(X2,X3) = k4_xxreal_1(sK0,X1) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
158. ? [X1] : (? [X2] : (? [X3] : ((X1 != X3 | sK0 != X2) & ~r1_xxreal_0(X1,sK0) & k4_xxreal_1(X2,X3) = k4_xxreal_1(sK0,X1) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (? [X3] : ((sK1 != X3 | sK0 != X2) & ~r1_xxreal_0(sK1,sK0) & k4_xxreal_1(X2,X3) = k4_xxreal_1(sK0,sK1) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
159. ? [X2] : (? [X3] : ((sK1 != X3 | sK0 != X2) & ~r1_xxreal_0(sK1,sK0) & k4_xxreal_1(X2,X3) = k4_xxreal_1(sK0,sK1) & v1_xxreal_0(X3)) & v1_xxreal_0(X2)) => (? [X3] : ((sK1 != X3 | sK0 != sK2) & ~r1_xxreal_0(sK1,sK0) & k4_xxreal_1(sK0,sK1) = k4_xxreal_1(sK2,X3) & v1_xxreal_0(X3)) & v1_xxreal_0(sK2)) [choice axiom]
160. ? [X3] : ((sK1 != X3 | sK0 != sK2) & ~r1_xxreal_0(sK1,sK0) & k4_xxreal_1(sK0,sK1) = k4_xxreal_1(sK2,X3) & v1_xxreal_0(X3)) => ((sK1 != sK3 | sK0 != sK2) & ~r1_xxreal_0(sK1,sK0) & k4_xxreal_1(sK0,sK1) = k4_xxreal_1(sK2,sK3) & v1_xxreal_0(sK3)) [choice axiom]
161. ((((sK1 != sK3 | sK0 != sK2) & ~r1_xxreal_0(sK1,sK0) & k4_xxreal_1(sK0,sK1) = k4_xxreal_1(sK2,sK3) & v1_xxreal_0(sK3)) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 97,160,159,158,157]
197. v1_xxreal_0(sK0) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
198. v1_xxreal_0(sK1) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
199. v1_xxreal_0(sK2) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
200. v1_xxreal_0(sK3) [cnf transformation 161] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
201. k4_xxreal_1(sK0,sK1) = k4_xxreal_1(sK2,sK3) [cnf transformation 161] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
202. ~r1_xxreal_0(sK1,sK0) [cnf transformation 161] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
203. sK1 != sK3 | sK0 != sK2 [cnf transformation 161] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
230. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 120] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
231. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 122] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
232. ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(X2,X3)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X2,X0) | ~v1_xxreal_0(X3) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 124] {a:0,w:21,wCS:8400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
233. ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(X2,X3)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X1,X3) | ~v1_xxreal_0(X3) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 124] {a:0,w:21,wCS:8400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
262. r1_tarski(X0,X0) [cnf transformation 86] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
266. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 154] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
299. 1 <=> sK0 = sK2 [avatar definition]
303. 2 <=> sK1 = sK3 [avatar definition]
306. ~1 | ~2 [avatar split clause 203,303,299]
308. 3 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
310. ~r1_xxreal_0(sK1,sK0) <- (~3) [avatar component clause 308] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
311. ~3 [avatar split clause 202,308]
313. 4 <=> k4_xxreal_1(sK0,sK1) = k4_xxreal_1(sK2,sK3) [avatar definition]
315. k4_xxreal_1(sK0,sK1) = k4_xxreal_1(sK2,sK3) <- (4) [avatar component clause 313] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
316. 4 [avatar split clause 201,313]
318. 5 <=> v1_xxreal_0(sK3) [avatar definition]
320. v1_xxreal_0(sK3) <- (5) [avatar component clause 318] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
321. 5 [avatar split clause 200,318]
323. 6 <=> v1_xxreal_0(sK2) [avatar definition]
325. v1_xxreal_0(sK2) <- (6) [avatar component clause 323] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
326. 6 [avatar split clause 199,323]
328. 7 <=> v1_xxreal_0(sK1) [avatar definition]
330. v1_xxreal_0(sK1) <- (7) [avatar component clause 328] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
331. 7 [avatar split clause 198,328]
333. 8 <=> v1_xxreal_0(sK0) [avatar definition]
335. v1_xxreal_0(sK0) <- (8) [avatar component clause 333] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
336. 8 [avatar split clause 197,333]
841. r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~3) [resolution 266,310] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
844. r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK0) <- (~3, 7) [subsumption resolution 841,330] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
845. r1_xxreal_0(sK0,sK1) <- (~3, 7, 8) [subsumption resolution 844,335] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
847. 92 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
849. r1_xxreal_0(sK0,sK1) <- (92) [avatar component clause 847] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
850. 92 | 3 | ~7 | ~8 [avatar split clause 845,333,328,308,847]
912. ~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) <- (4) [superposition 232,315] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
913. ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(sK2,X0) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4) [superposition 232,315] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
915. ~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK2) <- (4, 5) [subsumption resolution 912,320] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
916. ~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5, 6) [subsumption resolution 915,325] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
918. 99 <=> r1_xxreal_0(sK3,sK2) [avatar definition]
920. r1_xxreal_0(sK3,sK2) <- (99) [avatar component clause 918] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
922. 100 <=> ! [X1,X0] : (~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | r1_xxreal_0(X0,sK2)) [avatar definition]
923. ~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | r1_xxreal_0(X0,sK2) <- (100) [avatar component clause 922] {a:1,w:14,wCS:1400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
924. 99 | 100 | ~4 | ~5 | ~6 [avatar split clause 916,323,318,313,922,918]
925. ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(sK2,X0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5) [subsumption resolution 913,320] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
926. ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(sK2,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5, 6) [subsumption resolution 925,325] {a:1,w:17,wCS:1700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
929. ~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(sK3,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) <- (4) [superposition 233,315] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
930. ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X1,sK3) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4) [superposition 233,315] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
932. ~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(sK3,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK2) <- (4, 5) [subsumption resolution 929,320] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
933. ~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | r1_xxreal_0(sK3,sK2) | r1_xxreal_0(sK3,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5, 6) [subsumption resolution 932,325] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
935. 101 <=> ! [X1,X0] : (~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | r1_xxreal_0(sK3,X1)) [avatar definition]
936. ~r1_tarski(k4_xxreal_1(sK0,sK1),k4_xxreal_1(X0,X1)) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | r1_xxreal_0(sK3,X1) <- (101) [avatar component clause 935] {a:1,w:14,wCS:1400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
937. 99 | 101 | ~4 | ~5 | ~6 [avatar split clause 933,323,318,313,935,918]
938. ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X1,sK3) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5) [subsumption resolution 930,320] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
939. ~r1_tarski(k4_xxreal_1(X0,X1),k4_xxreal_1(sK0,sK1)) | r1_xxreal_0(X1,X0) | r1_xxreal_0(X1,sK3) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) <- (4, 5, 6) [subsumption resolution 938,325] {a:1,w:17,wCS:1700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
940. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(X0) <- (92) [resolution 849,231] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
942. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(X0) <- (7, 92) [subsumption resolution 940,330] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
943. ~r1_xxreal_0(X0,sK0) | r1_xxreal_0(X0,sK1) | ~v1_xxreal_0(X0) <- (7, 8, 92) [subsumption resolution 942,335] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
945. sK2 = sK3 | ~r1_xxreal_0(sK2,sK3) | ~v1_xxreal_0(sK3) | ~v1_xxreal_0(sK2) <- (99) [resolution 920,230] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
948. sK2 = sK3 | ~r1_xxreal_0(sK2,sK3) | ~v1_xxreal_0(sK2) <- (5, 99) [subsumption resolution 945,320] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
949. sK2 = sK3 | ~r1_xxreal_0(sK2,sK3) <- (5, 6, 99) [subsumption resolution 948,325] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
951. 102 <=> r1_xxreal_0(sK2,sK3) [avatar definition]
953. ~r1_xxreal_0(sK2,sK3) <- (~102) [avatar component clause 951] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
955. 103 <=> sK2 = sK3 [avatar definition]
958. ~102 | 103 | ~5 | ~6 | ~99 [avatar split clause 949,918,323,318,955,951]
986. r1_xxreal_0(sK1,sK0) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4, 5, 6) [resolution 926,262] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
988. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~3, 4, 5, 6) [subsumption resolution 986,310] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
989. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) <- (~3, 4, 5, 6, 7) [subsumption resolution 988,330] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
990. r1_xxreal_0(sK2,sK0) <- (~3, 4, 5, 6, 7, 8) [subsumption resolution 989,335] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
992. 108 <=> r1_xxreal_0(sK2,sK0) [avatar definition]
994. r1_xxreal_0(sK2,sK0) <- (108) [avatar component clause 992] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
995. 108 | 3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 990,333,328,323,318,313,308,992]
996. r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK2) <- (7, 8, 92, 108) [resolution 994,943] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
998. sK0 = sK2 | ~r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (108) [resolution 994,230] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
999. r1_xxreal_0(sK2,sK1) <- (6, 7, 8, 92, 108) [subsumption resolution 996,325] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
1001. 109 <=> r1_xxreal_0(sK2,sK1) [avatar definition]
1003. r1_xxreal_0(sK2,sK1) <- (109) [avatar component clause 1001] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1004. 109 | ~6 | ~7 | ~8 | ~92 | ~108 [avatar split clause 999,992,847,333,328,323,1001]
1011. 110 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
1015. r1_xxreal_0(sK1,sK0) | r1_xxreal_0(sK1,sK3) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4, 5, 6) [resolution 939,262] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1017. r1_xxreal_0(sK1,sK3) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~3, 4, 5, 6) [subsumption resolution 1015,310] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1018. r1_xxreal_0(sK1,sK3) | ~v1_xxreal_0(sK0) <- (~3, 4, 5, 6, 7) [subsumption resolution 1017,330] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1019. r1_xxreal_0(sK1,sK3) <- (~3, 4, 5, 6, 7, 8) [subsumption resolution 1018,335] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1021. 111 <=> r1_xxreal_0(sK1,sK3) [avatar definition]
1023. r1_xxreal_0(sK1,sK3) <- (111) [avatar component clause 1021] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1024. 111 | 3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 1019,333,328,323,318,313,308,1021]
1026. sK1 = sK3 | ~r1_xxreal_0(sK3,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK3) <- (111) [resolution 1023,230] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1029. sK1 = sK3 | ~r1_xxreal_0(sK3,sK1) | ~v1_xxreal_0(sK3) <- (7, 111) [subsumption resolution 1026,330] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1030. sK1 = sK3 | ~r1_xxreal_0(sK3,sK1) <- (5, 7, 111) [subsumption resolution 1029,320] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1032. 112 <=> r1_xxreal_0(sK3,sK1) [avatar definition]
1035. ~112 | 2 | ~5 | ~7 | ~111 [avatar split clause 1030,1021,328,318,303,1032]
1038. $false <- (5, 6, 7, ~102, 109, 111) [unit resulting resolution 325,330,320,1023,953,1003,231] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:44,thDist:-44}
1041. ~5 | ~6 | ~7 | 102 | ~109 | ~111 [avatar contradiction clause 1038]
1063. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | r1_xxreal_0(sK0,sK2) <- (100) [resolution 923,262] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1065. ~v1_xxreal_0(sK1) | r1_xxreal_0(sK0,sK2) <- (8, 100) [subsumption resolution 1063,335] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1066. r1_xxreal_0(sK0,sK2) <- (7, 8, 100) [subsumption resolution 1065,330] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1067. 110 | ~7 | ~8 | ~100 [avatar split clause 1066,922,333,328,1011]
1068. sK2 != sK3 | r1_xxreal_0(sK3,sK1) | ~r1_xxreal_0(sK2,sK1) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
1069. sK0 = sK2 | ~r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK0) <- (6, 108) [subsumption resolution 998,325] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1070. sK0 = sK2 | ~r1_xxreal_0(sK0,sK2) <- (6, 8, 108) [subsumption resolution 1069,335] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1071. ~110 | 1 | ~6 | ~8 | ~108 [avatar split clause 1070,992,333,323,299,1011]
1102. sK1 != sK3 | sK2 != sK3 | r1_xxreal_0(sK1,sK0) | ~r1_xxreal_0(sK2,sK0) [theory tautology sat conflict] {a:0,w:12,wCS:4800,sine:255,thAx:0,allAx:1,thDist:-1}
1130. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | r1_xxreal_0(sK3,sK1) <- (101) [resolution 936,262] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1132. ~v1_xxreal_0(sK1) | r1_xxreal_0(sK3,sK1) <- (8, 101) [subsumption resolution 1130,335] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1133. r1_xxreal_0(sK3,sK1) <- (7, 8, 101) [subsumption resolution 1132,330] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1134. 112 | ~7 | ~8 | ~101 [avatar split clause 1133,935,333,328,1032]
1135. $false [avatar sat refutation 306,311,316,321,326,331,336,850,924,937,958,995,1004,1024,1035,1041,1067,1068,1071,1102,1134] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:768,thDist:-768}
% SZS output end Proof for xxreal_1__t67_xxreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 76
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 100
% Generated clauses: 468
% Active clauses: 150
% Passive clauses: 226
% Final active clauses: 130
% Final passive clauses: 69
% Main loop iterations started: 155
% 
% >>> Simplifying Inferences
% Duplicate literals: 28
% Fw subsumption resolutions: 65
% Fw demodulations: 1
% Bw demodulations: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 1
% Forward subsumptions: 42
% 
% >>> Generating Inferences
% Binary resolution: 133
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 14
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 19
% Split components: 45
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 115
% SAT solver unit clauses: 38
% SAT solver binary clauses: 36
% 
% Memory used [KB]: 5628
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% non unit clause index maintenance: 0.001 s
% resolution: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

