i: [1,47,0,2,1,1,0,0] t81_rfunct_1
a: [3,85,0,2,1,47]
s: 85
i: [1,46,0,2,1,1,0,0] t81_rfunct_1
a: [3,90,0,2,1,46]
s: 90
i: [1,45,0,2,1,1,0,0] t81_rfunct_1
a: [3,95,0,2,1,45]
s: 95
i: [1,48,0,5,1,1,0,0] t81_rfunct_1
a: [3,80,0,5,1,48]
s: 80
i: [1,76,0,2,1,0,0,1] rc1_funct_1
a: [3,100,0,2,1,76]
s: 100
i: [1,75,0,2,1,0,0,1] rc1_funct_1
a: [3,105,0,2,1,75]
s: 105
i: [1,50,0,6,3,0,0,255] cc5_valued_0
s: 50
i: [1,52,0,6,3,0,0,2] cc6_valued_0
s: 52
d: [2,107,1,4,2,1,40,50,85]
d: [2,108,1,2,1,2,21,107,95]
a: [3,112,1,2,1,108]
s: 112
i: [1,61,0,6,2,0,0,1] dt_k5_relat_1
s: 61
i: [1,63,0,8,3,0,0,255] fc44_valued_0
s: 63
d: [2,114,1,4,2,1,40,52,85]
d: [2,115,1,2,1,2,21,114,95]
a: [3,119,1,2,1,115]
s: 119
i: [1,64,0,8,3,0,0,255] fc44_valued_0
s: 64
i: [1,65,0,8,3,0,0,2] fc43_valued_0
s: 65
i: [1,66,0,8,3,0,0,2] fc43_valued_0
s: 66
i: [1,67,0,8,3,0,0,1] fc45_valued_0
s: 67
i: [1,68,0,8,3,0,0,1] fc45_valued_0
s: 68
i: [1,73,0,8,3,0,0,1] fc8_funct_1
s: 73
i: [1,74,0,8,3,0,0,1] fc8_funct_1
s: 74
i: [1,54,0,9,4,0,0,1] dt_k54_valued_1
s: 54
i: [1,55,0,9,4,0,0,1] dt_k54_valued_1
s: 55
d: [2,121,1,12,4,0,40,68,52]
d: [2,123,1,8,3,0,21,121,61]
s: 123
i: [1,56,0,9,4,0,0,1] dt_k54_valued_1
s: 56
d: [2,122,1,12,4,0,40,68,50]
d: [2,124,1,8,3,0,21,122,61]
s: 124
i: [1,57,0,9,4,0,0,1] fc12_rfunct_1
s: 57
i: [1,58,0,9,4,0,0,1] fc12_rfunct_1
s: 58
i: [1,59,0,9,4,0,0,1] fc12_rfunct_1
s: 59
i: [1,60,0,9,4,0,0,1] fc12_rfunct_1
s: 60
d: [2,125,1,12,5,0,40,56,52]
d: [2,127,1,9,4,0,21,125,54]
s: 127
i: [1,62,0,11,2,0,0,255] rd5_relat_1
s: 62
i: [1,53,0,12,4,0,0,255] projectivity_k54_valued_1
s: 53
d: [2,126,1,12,5,0,40,56,50]
d: [2,128,1,9,4,0,21,126,54]
s: 128
i: [1,69,0,12,5,0,0,1] fc7_rfunct_1
s: 69
d: [2,129,1,12,5,0,40,59,52]
d: [2,131,1,9,4,0,21,129,57]
s: 131
i: [1,70,0,12,5,0,0,1] fc7_rfunct_1
s: 70
i: [1,71,0,12,5,0,0,1] fc7_rfunct_1
s: 71
d: [2,130,1,12,5,0,40,59,50]
d: [2,132,1,9,4,0,21,130,57]
s: 132
i: [1,72,0,12,5,0,0,1] fc7_rfunct_1
s: 72
d: [2,156,1,12,4,1,40,72,80]
a: [3,161,1,3,1,156]
s: 161
d: [2,175,2,0,0,4,52,95,90,85,161,57]
e: 175
d: [2,176,2,0,0,4,52,95,90,119,161,54]
e: 176
d: [2,177,2,6,3,1,40,161,57]
d: [2,181,2,4,2,2,21,177,85]
d: [2,182,2,2,1,3,21,181,90]
d: [2,183,2,0,0,4,21,182,95]
e: 183
d: [2,178,2,6,3,1,40,161,54]
d: [2,185,2,4,2,2,21,178,119]
d: [2,186,2,2,1,3,21,185,90]
d: [2,187,2,0,0,4,21,186,95]
e: 187
r: 161
a: [3,160,1,3,1,156]
s: 160
a: [3,165,1,3,1,156]
s: 165
d: [2,189,2,0,0,4,52,95,90,85,165,58]
e: 189
d: [2,190,2,0,0,4,52,95,90,119,165,55]
e: 190
d: [2,191,2,6,3,1,40,165,58]
d: [2,195,2,4,2,2,21,191,85]
d: [2,196,2,2,1,3,21,195,90]
d: [2,197,2,0,0,4,21,196,95]
e: 197
d: [2,192,2,6,3,1,40,165,55]
d: [2,199,2,4,2,2,21,192,119]
d: [2,200,2,2,1,3,21,199,90]
d: [2,201,2,0,0,4,21,200,95]
e: 201
r: 165
a: [3,164,1,3,1,156]
s: 164
a: [3,169,1,3,1,156]
s: 169
d: [2,203,2,0,0,4,52,95,90,85,169,59]
e: 203
d: [2,204,2,0,0,4,52,95,90,119,169,56]
e: 204
d: [2,205,2,6,3,1,40,169,59]
d: [2,209,2,4,2,2,21,205,85]
d: [2,210,2,2,1,3,21,209,90]
d: [2,211,2,0,0,4,21,210,95]
e: 211
d: [2,206,2,6,3,1,40,169,56]
d: [2,213,2,4,2,2,21,206,119]
d: [2,214,2,2,1,3,21,213,90]
d: [2,215,2,0,0,4,21,214,95]
e: 215
r: 169
a: [3,168,1,3,1,156]
s: 168
a: [3,173,1,3,1,156]
s: 173
d: [2,231,2,0,0,4,52,95,90,85,173,60]
e: 231
d: [2,232,2,6,3,1,40,173,60]
d: [2,234,2,4,2,2,21,232,85]
d: [2,235,2,2,1,3,21,234,90]
d: [2,236,2,0,0,4,21,235,95]
e: 236
f: 236,211,197,183,156,45,46,47,48
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rfunct_1__t81_rfunct_1
% SZS output start Proof for rfunct_1__t81_rfunct_1
1. ! [X0] : ! [X1] : ((v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => v2_seq_2(k5_relat_1(k54_valued_1(X1),X0))) [input t81_rfunct_1]
2. ~! [X0] : ! [X1] : ((v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => v2_seq_2(k5_relat_1(k54_valued_1(X1),X0))) [negated conjecture 1]
10. ! [X0] : ((v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v2_seq_2(k54_valued_1(X0)) & v3_valued_0(k54_valued_1(X0)) & v1_funct_1(k54_valued_1(X0)) & v1_relat_1(k54_valued_1(X0)))) [input fc12_rfunct_1]
12. ! [X0,X1] : ((v2_seq_2(X0) & v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v2_seq_2(k5_relat_1(X0,X1)) & v3_valued_0(k5_relat_1(X0,X1)) & v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1)))) [input fc7_rfunct_1]
16. ~! [X0,X1] : ((v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => v2_seq_2(k5_relat_1(k54_valued_1(X1),X0))) [flattening 2]
17. ? [X0,X1] : (~v2_seq_2(k5_relat_1(k54_valued_1(X1),X0)) & (v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1))) [ennf transformation 16]
18. ? [X0,X1] : (~v2_seq_2(k5_relat_1(k54_valued_1(X1),X0)) & v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) [flattening 17]
27. ! [X0] : ((v2_seq_2(k54_valued_1(X0)) & v3_valued_0(k54_valued_1(X0)) & v1_funct_1(k54_valued_1(X0)) & v1_relat_1(k54_valued_1(X0))) | (~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 10]
28. ! [X0] : ((v2_seq_2(k54_valued_1(X0)) & v3_valued_0(k54_valued_1(X0)) & v1_funct_1(k54_valued_1(X0)) & v1_relat_1(k54_valued_1(X0))) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 27]
37. ! [X0,X1] : ((v2_seq_2(k5_relat_1(X0,X1)) & v3_valued_0(k5_relat_1(X0,X1)) & v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | (~v2_seq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 12]
38. ! [X0,X1] : ((v2_seq_2(k5_relat_1(X0,X1)) & v3_valued_0(k5_relat_1(X0,X1)) & v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | ~v2_seq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 37]
41. ? [X0,X1] : (~v2_seq_2(k5_relat_1(k54_valued_1(X1),X0)) & v3_valued_0(X1) & v1_funct_1(X1) & v1_relat_1(X1)) => (~v2_seq_2(k5_relat_1(k54_valued_1(sK1),sK0)) & v3_valued_0(sK1) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
42. ~v2_seq_2(k5_relat_1(k54_valued_1(sK1),sK0)) & v3_valued_0(sK1) & v1_funct_1(sK1) & v1_relat_1(sK1) [skolemisation 18,41]
45. v1_relat_1(sK1) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. v1_funct_1(sK1) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. v3_valued_0(sK1) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. ~v2_seq_2(k5_relat_1(k54_valued_1(sK1),sK0)) [cnf transformation 42] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
57. v1_relat_1(k54_valued_1(X0)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. v1_funct_1(k54_valued_1(X0)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. v3_valued_0(k54_valued_1(X0)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. v2_seq_2(k54_valued_1(X0)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. v2_seq_2(k5_relat_1(X0,X1)) | ~v2_seq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 38] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
78. 1 <=> v2_seq_2(k5_relat_1(k54_valued_1(sK1),sK0)) [avatar definition]
80. ~v2_seq_2(k5_relat_1(k54_valued_1(sK1),sK0)) <- (~1) [avatar component clause 78] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
81. ~1 [avatar split clause 48,78]
83. 2 <=> v3_valued_0(sK1) [avatar definition]
85. v3_valued_0(sK1) <- (2) [avatar component clause 83] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
86. 2 [avatar split clause 47,83]
88. 3 <=> v1_funct_1(sK1) [avatar definition]
90. v1_funct_1(sK1) <- (3) [avatar component clause 88] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
91. 3 [avatar split clause 46,88]
93. 4 <=> v1_relat_1(sK1) [avatar definition]
95. v1_relat_1(sK1) <- (4) [avatar component clause 93] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
96. 4 [avatar split clause 45,93]
156. ~v2_seq_2(k54_valued_1(sK1)) | ~v3_valued_0(k54_valued_1(sK1)) | ~v1_funct_1(k54_valued_1(sK1)) | ~v1_relat_1(k54_valued_1(sK1)) <- (~1) [resolution 72,80] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
159. 9 <=> v1_relat_1(k54_valued_1(sK1)) [avatar definition]
161. ~v1_relat_1(k54_valued_1(sK1)) <- (~9) [avatar component clause 159] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
163. 10 <=> v1_funct_1(k54_valued_1(sK1)) [avatar definition]
165. ~v1_funct_1(k54_valued_1(sK1)) <- (~10) [avatar component clause 163] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
167. 11 <=> v3_valued_0(k54_valued_1(sK1)) [avatar definition]
169. ~v3_valued_0(k54_valued_1(sK1)) <- (~11) [avatar component clause 167] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
171. 12 <=> v2_seq_2(k54_valued_1(sK1)) [avatar definition]
173. ~v2_seq_2(k54_valued_1(sK1)) <- (~12) [avatar component clause 171] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
174. ~9 | ~10 | ~11 | ~12 | 1 [avatar split clause 156,78,171,167,163,159]
177. ~v3_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~9) [resolution 161,57] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
181. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, ~9) [subsumption resolution 177,85] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
182. ~v1_relat_1(sK1) <- (2, 3, ~9) [subsumption resolution 181,90] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
183. $false <- (2, 3, 4, ~9) [subsumption resolution 182,95] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
184. ~2 | ~3 | ~4 | 9 [avatar contradiction clause 183]
191. ~v3_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~10) [resolution 165,58] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
195. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, ~10) [subsumption resolution 191,85] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
196. ~v1_relat_1(sK1) <- (2, 3, ~10) [subsumption resolution 195,90] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
197. $false <- (2, 3, 4, ~10) [subsumption resolution 196,95] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
198. ~2 | ~3 | ~4 | 10 [avatar contradiction clause 197]
205. ~v3_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~11) [resolution 169,59] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
209. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, ~11) [subsumption resolution 205,85] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
210. ~v1_relat_1(sK1) <- (2, 3, ~11) [subsumption resolution 209,90] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
211. $false <- (2, 3, 4, ~11) [subsumption resolution 210,95] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
212. ~2 | ~3 | ~4 | 11 [avatar contradiction clause 211]
232. ~v3_valued_0(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~12) [resolution 173,60] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
234. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, ~12) [subsumption resolution 232,85] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
235. ~v1_relat_1(sK1) <- (2, 3, ~12) [subsumption resolution 234,90] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
236. $false <- (2, 3, 4, ~12) [subsumption resolution 235,95] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
237. ~2 | ~3 | ~4 | 12 [avatar contradiction clause 236]
238. $false [avatar sat refutation 81,86,91,96,174,184,198,212,237] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:153,thDist:-153}
% SZS output end Proof for rfunct_1__t81_rfunct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 32
% Generated clauses: 129
% Active clauses: 38
% Passive clauses: 48
% Final active clauses: 35
% Final passive clauses: 3
% Main loop iterations started: 45
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 34
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 15
% 
% >>> Generating Inferences
% Binary resolution: 18
% Unit resulting resolution: 7
% Forward superposition: 4
% Backward superposition: 15
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 4
% Sat splitting refutations: 14
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 27
% SAT solver unit clauses: 6
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

