i: [1,239,0,2,1,1,0,0] t126_xreal_1
a: [3,379,0,2,1,239]
s: 379
i: [1,238,0,2,1,1,0,0] t126_xreal_1
a: [3,384,0,2,1,238]
s: 384
i: [1,241,0,3,1,1,0,0] t126_xreal_1
a: [3,369,0,3,1,241]
s: 369
i: [1,240,0,3,1,1,0,0] t126_xreal_1
a: [3,374,0,3,1,240]
s: 374
i: [1,242,0,5,1,1,0,0] t126_xreal_1
a: [3,364,0,5,1,242]
s: 364
i: [1,243,0,2,1,0,0,255] fc1_numbers
a: [3,389,0,2,1,243]
s: 389
i: [1,244,0,2,1,0,0,255] spc1_boole
a: [3,394,0,2,1,244]
s: 394
i: [1,245,0,2,1,0,0,255] fc1_xboole_0
a: [3,399,0,2,1,245]
s: 399
i: [1,254,0,2,1,0,0,255] fc6_ordinal1
a: [3,439,0,2,1,254]
s: 439
i: [1,253,0,2,1,0,0,255] fc6_ordinal1
a: [3,444,0,2,1,253]
s: 444
i: [1,255,0,2,1,0,0,255] spc1_numerals
a: [3,464,0,2,1,255]
s: 464
i: [1,328,0,2,1,0,0,2] rc2_xboole_0
a: [3,476,0,2,1,328]
s: 476
i: [1,329,0,2,1,0,0,255] rc4_ordinal1
a: [3,481,0,2,1,329]
s: 481
i: [1,330,0,2,1,0,0,255] rc1_ordinal1
a: [3,486,0,2,1,330]
s: 486
i: [1,331,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,491,0,2,1,331]
s: 491
i: [1,332,0,2,1,0,0,1] rc1_xreal_0
a: [3,496,0,2,1,332]
s: 496
i: [1,333,0,2,1,0,0,2] rc1_xxreal_0
a: [3,501,0,2,1,333]
s: 501
i: [1,334,0,2,1,0,0,2] rc1_xboole_0
a: [3,506,0,2,1,334]
s: 506
i: [1,338,0,2,1,0,0,255] rc2_ordinal1
a: [3,511,0,2,1,338]
s: 511
i: [1,337,0,2,1,0,0,255] rc2_ordinal1
a: [3,516,0,2,1,337]
s: 516
i: [1,336,0,2,1,0,0,255] rc2_ordinal1
a: [3,521,0,2,1,336]
s: 521
i: [1,335,0,2,1,0,0,255] rc2_ordinal1
a: [3,526,0,2,1,335]
s: 526
i: [1,340,0,2,1,0,0,255] rc5_ordinal1
a: [3,531,0,2,1,340]
s: 531
i: [1,339,0,2,1,0,0,255] rc5_ordinal1
a: [3,536,0,2,1,339]
s: 536
i: [1,342,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,541,0,2,1,342]
s: 541
i: [1,341,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,546,0,2,1,341]
s: 546
i: [1,346,0,2,1,0,0,2] rc3_xreal_0
a: [3,551,0,2,1,346]
s: 551
i: [1,345,0,2,1,0,0,2] rc3_xreal_0
a: [3,556,0,2,1,345]
s: 556
i: [1,344,0,2,1,0,0,2] rc3_xreal_0
a: [3,561,0,2,1,344]
s: 561
i: [1,343,0,2,1,0,0,2] rc3_xreal_0
a: [3,566,0,2,1,343]
s: 566
i: [1,350,0,2,1,0,0,2] rc2_xreal_0
a: [3,571,0,2,1,350]
s: 571
i: [1,349,0,2,1,0,0,2] rc2_xreal_0
a: [3,576,0,2,1,349]
s: 576
i: [1,348,0,2,1,0,0,2] rc2_xreal_0
a: [3,581,0,2,1,348]
s: 581
i: [1,347,0,2,1,0,0,2] rc2_xreal_0
a: [3,586,0,2,1,347]
s: 586
i: [1,352,0,2,1,0,0,2] rc3_xxreal_0
a: [3,591,0,2,1,352]
s: 591
i: [1,351,0,2,1,0,0,2] rc3_xxreal_0
a: [3,596,0,2,1,351]
s: 596
i: [1,354,0,2,1,0,0,2] rc2_xxreal_0
a: [3,601,0,2,1,354]
s: 601
i: [1,353,0,2,1,0,0,2] rc2_xxreal_0
a: [3,606,0,2,1,353]
s: 606
i: [1,358,0,2,1,0,0,2] rc4_xreal_0
a: [3,611,0,2,1,358]
s: 611
i: [1,357,0,2,1,0,0,2] rc4_xreal_0
a: [3,616,0,2,1,357]
s: 616
i: [1,356,0,2,1,0,0,2] rc4_xreal_0
a: [3,621,0,2,1,356]
s: 621
i: [1,355,0,2,1,0,0,2] rc4_xreal_0
a: [3,626,0,2,1,355]
s: 626
i: [1,360,0,2,1,0,0,2] rc4_xxreal_0
a: [3,631,0,2,1,360]
s: 631
i: [1,359,0,2,1,0,0,2] rc4_xxreal_0
a: [3,636,0,2,1,359]
s: 636
i: [1,311,0,3,1,0,0,255] reflexivity_r1_tarski
s: 311
i: [1,246,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,404,0,3,1,246]
s: 404
i: [1,247,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,409,0,3,1,247]
s: 409
r: 399
d: [2,638,0,2,1,2,23,399,409]
a: [3,642,0,2,1,638]
s: 642
i: [1,248,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,414,0,3,1,248]
s: 414
i: [1,249,0,3,1,0,0,255] t1_numerals
a: [3,419,0,3,1,249]
s: 419
i: [1,258,0,3,1,0,0,255] spc1_numerals
a: [3,449,0,3,1,258]
s: 449
i: [1,257,0,3,1,0,0,255] spc1_numerals
a: [3,454,0,3,1,257]
s: 454
d: [2,644,0,3,1,2,22,419,409]
a: [3,648,0,3,1,644]
s: 648
d: [2,650,0,3,1,2,22,454,404]
a: [3,654,0,3,1,650]
s: 654
i: [1,260,0,4,2,0,0,255] cc6_ordinal1
s: 260
i: [1,261,0,4,2,0,0,255] cc2_xcmplx_0
s: 261
d: [2,656,1,2,1,1,40,260,481]
a: [3,661,1,2,1,656]
s: 661
i: [1,262,0,4,2,0,0,255] cc2_xreal_0
s: 262
d: [2,657,1,2,1,1,40,260,531]
a: [3,666,1,2,1,657]
s: 666
i: [1,263,0,4,2,0,0,255] cc2_xxreal_0
s: 263
i: [1,264,0,4,2,0,0,255] cc1_ordinal1
s: 264
d: [2,668,1,2,1,1,40,261,481]
a: [3,673,1,2,1,668]
s: 673
i: [1,265,0,4,2,0,0,255] cc1_ordinal1
s: 265
d: [2,669,1,2,1,1,40,261,531]
a: [3,678,1,2,1,669]
s: 678
i: [1,269,0,4,2,0,0,2] cc3_xreal_0
s: 269
i: [1,270,0,4,2,0,0,2] cc4_xreal_0
s: 270
d: [2,754,1,2,1,1,40,269,384]
a: [3,763,1,2,1,754]
s: 763
i: [1,280,0,4,2,0,0,255] cc9_ordinal1
s: 280
d: [2,755,1,2,1,1,40,269,379]
a: [3,768,1,2,1,755]
s: 768
i: [1,281,0,4,2,0,0,255] cc7_ordinal1
s: 281
i: [1,282,0,4,2,0,0,255] cc3_ordinal1
s: 282
d: [2,775,1,2,1,1,40,270,384]
a: [3,784,1,2,1,775]
s: 784
i: [1,310,0,4,1,0,0,255] existence_m1_subset_1
s: 310
d: [2,776,1,2,1,1,40,270,379]
a: [3,789,1,2,1,776]
s: 789
i: [1,250,0,4,1,0,0,255] dt_k5_numbers
a: [3,424,0,4,1,250]
s: 424
i: [1,251,0,4,1,0,0,255] dt_k6_numbers
a: [3,429,0,4,1,251]
s: 429
d: [2,680,1,2,1,1,40,262,481]
a: [3,685,1,2,1,680]
s: 685
i: [1,256,0,4,1,0,0,255] spc1_numerals
a: [3,459,0,4,1,256]
s: 459
d: [2,681,1,2,1,1,40,262,531]
a: [3,690,1,2,1,681]
s: 690
d: [2,868,0,4,1,2,22,424,404]
a: [3,872,0,4,1,868]
s: 872
d: [2,874,0,4,1,2,22,429,404]
a: [3,878,0,4,1,874]
s: 878
d: [2,692,1,2,1,1,40,263,481]
a: [3,697,1,2,1,692]
s: 697
d: [2,883,0,4,1,2,22,459,404]
a: [3,887,0,4,1,883]
s: 887
d: [2,693,1,2,1,1,40,263,531]
a: [3,702,1,2,1,693]
s: 702
i: [1,283,0,5,2,0,0,255] t6_boole
s: 283
i: [1,284,0,5,2,0,0,255] cc8_ordinal1
s: 284
d: [2,704,1,2,1,1,40,264,439]
a: [3,712,1,2,1,704]
s: 712
i: [1,285,0,5,2,0,0,255] cc1_xcmplx_0
s: 285
d: [2,706,1,2,1,1,40,264,486]
a: [3,722,1,2,1,706]
s: 722
i: [1,286,0,5,2,0,0,255] cc1_xreal_0
s: 286
i: [1,307,0,5,2,0,0,255] t7_boole
s: 307
d: [2,729,1,2,1,1,40,265,439]
a: [3,737,1,2,1,729]
s: 737
i: [1,252,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,434,0,5,1,252]
s: 434
d: [2,731,1,2,1,1,40,265,486]
a: [3,747,1,2,1,731]
s: 747
i: [1,325,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,471,0,5,2,325]
s: 471
d: [2,892,0,5,2,1,22,283,409]
s: 892
d: [2,935,1,5,2,2,44,374,892]
s: 935
i: [1,291,0,6,3,0,0,255] cc2_ordinal1
s: 291
d: [2,936,1,5,2,2,44,369,892]
s: 936
i: [1,294,0,6,3,0,0,2] cc5_xxreal_0
s: 294
i: [1,296,0,6,3,0,0,2] cc5_xxreal_0
s: 296
d: [2,756,1,2,1,1,40,269,496]
a: [3,773,1,2,1,756]
s: 773
i: [1,297,0,6,3,0,0,2] cc3_xxreal_0
s: 297
d: [2,777,1,2,1,1,40,270,496]
a: [3,794,1,2,1,777]
s: 794
i: [1,299,0,6,3,0,0,2] cc3_xxreal_0
s: 299
i: [1,301,0,6,3,0,0,2] cc7_xxreal_0
s: 301
d: [2,796,1,2,1,1,40,280,642]
a: [3,803,1,2,1,796]
s: 803
i: [1,302,0,6,3,0,0,2] cc7_xxreal_0
s: 302
d: [2,797,1,2,1,1,40,280,506]
a: [3,808,1,2,1,797]
s: 808
i: [1,303,0,6,2,0,0,255] t1_subset
s: 303
i: [1,312,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 312
d: [2,798,1,2,1,1,40,280,626]
a: [3,813,1,2,1,798]
s: 813
i: [1,266,0,7,3,0,0,255] cc5_ordinal1
s: 266
d: [2,799,1,2,1,1,40,280,636]
a: [3,818,1,2,1,799]
s: 818
i: [1,267,0,7,2,0,0,1] t2_arithm
s: 267
i: [1,268,0,7,2,0,0,255] t3_arithm
s: 268
d: [2,820,1,2,1,1,40,281,642]
a: [3,827,1,2,1,820]
s: 827
i: [1,305,0,7,2,0,0,255] t3_subset
s: 305
d: [2,821,1,2,1,1,40,281,506]
a: [3,832,1,2,1,821]
s: 832
i: [1,306,0,7,3,0,0,1] t8_boole
s: 306
i: [1,259,0,8,3,0,0,255] cc10_ordinal1
s: 259
d: [2,822,1,2,1,1,40,281,626]
a: [3,837,1,2,1,822]
s: 837
i: [1,288,0,8,4,0,0,2] cc4_xxreal_0
s: 288
d: [2,1062,2,6,3,1,40,288,784]
a: [3,1077,2,2,1,1062]
s: 1077
r: 1077
i: [1,290,0,8,4,0,0,2] cc6_xxreal_0
s: 290
i: [1,292,0,8,4,0,0,2] cc8_xxreal_0
s: 292
a: [3,1076,2,2,1,1062]
s: 1076
i: [1,304,0,8,3,0,0,255] t2_subset
s: 304
a: [3,1081,2,2,1,1062]
s: 1081
i: [1,322,0,8,3,0,0,1] fc3_xcmplx_0
s: 322
i: [1,324,0,8,3,0,0,1] fc6_xreal_0
s: 324
d: [2,1063,2,6,3,1,40,288,789]
a: [3,1090,2,2,1,1063]
s: 1090
r: 1090
i: [1,309,0,9,3,0,0,255] t5_subset
s: 309
a: [3,1089,2,2,1,1063]
s: 1089
i: [1,308,0,10,3,0,0,255] t4_subset
s: 308
i: [1,326,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 326
a: [3,1094,2,2,1,1063]
s: 1094
i: [1,273,0,11,5,0,0,1] t2_real
s: 273
d: [2,1243,1,8,4,1,40,273,369]
d: [2,1256,1,6,3,2,21,1243,384]
a: [3,1260,1,2,1,1256]
s: 1260
r: 1260
i: [1,274,0,11,5,0,0,1] t1_real
s: 274
r: 1076
r: 1081
i: [1,275,0,11,5,0,0,1] t8_real
s: 275
a: [3,1259,1,2,1,1256]
s: 1259
i: [1,276,0,11,5,0,0,1] t7_real
s: 276
s: 1077
i: [1,277,0,11,5,0,0,1] t3_real
s: 277
i: [1,278,0,11,5,0,0,1] t4_real
s: 278
s: 1081
d: [2,1320,3,0,0,3,52,784,1077,1081,297]
e: 1320
d: [2,1321,3,4,2,1,40,1081,297]
d: [2,1323,3,2,1,2,21,1321,1077]
d: [2,1324,3,0,0,3,21,1323,784]
e: 1324
r: 1077
i: [1,323,0,11,3,0,0,1] commutativity_k3_xcmplx_0
s: 323
s: 1076
i: [1,313,0,12,5,0,0,1] fc26_xreal_0
s: 313
i: [1,314,0,12,5,0,0,1] fc24_xreal_0
s: 314
a: [3,1084,2,2,1,1062]
s: 1084
i: [1,315,0,12,5,0,0,1] fc23_xreal_0
s: 315
r: 1076
r: 1081
s: 1077
i: [1,316,0,12,5,0,0,1] fc25_xreal_0
s: 316
i: [1,321,0,12,5,0,0,1] fc8_xcmplx_0
s: 321
a: [3,1080,2,2,1,1062]
s: 1080
i: [1,271,0,13,6,0,0,1] t5_real
s: 271
d: [2,1269,2,4,2,1,40,1260,296]
a: [3,1275,2,2,1,1269]
s: 1275
i: [1,272,0,13,6,0,0,1] t6_real
s: 272
i: [1,279,0,13,5,0,0,1] l128_xreal_1
s: 279
d: [2,1445,3,0,0,5,52,379,1080,1084,384,364,279]
e: 1445
d: [2,1450,1,11,5,1,44,364,279]
d: [2,1468,1,8,4,1,19,1450]
d: [2,1475,1,6,3,2,21,1468,384]
d: [2,1476,1,4,2,3,21,1475,1084]
d: [2,1477,1,2,1,4,21,1476,1080]
d: [2,1478,1,0,0,5,21,1477,379]
e: 1478
r: 1077
r: 1080
r: 1275
s: 1076
i: [1,320,0,14,4,0,0,255] existence_m2_subset_1
s: 320
s: 1081
i: [1,317,0,15,5,0,0,255] dt_m2_subset_1
s: 317
i: [1,318,0,15,5,0,0,255] redefinition_m2_subset_1
s: 318
a: [3,1274,2,2,1,1269]
s: 1274
d: [2,893,1,2,1,1,40,284,648]
d: [2,1003,2,2,1,1,40,827,263]
d: [2,1004,2,2,1,1,40,827,262]
d: [2,1270,2,4,2,1,40,1260,294]
d: [2,1277,2,2,1,2,21,1270,642]
d: [2,1284,1,8,4,1,40,274,374]
d: [2,1300,1,6,3,2,21,1284,384]
d: [2,1500,1,2,1,4,21,1476,379]
d: [2,1533,3,4,2,1,40,1274,297]
d: [2,1535,3,2,1,2,21,1533,642]
f: 1535,1500,1300,1277,1256,1004,1003,893,644,638,249,247,245,238,239,240,241,242
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t126_xreal_1
% SZS output start Proof for xreal_1__t126_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,k6_numbers)) => k6_numbers = k3_xcmplx_0(X0,X1)))) [input t126_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,k6_numbers)) => k6_numbers = k3_xcmplx_0(X0,X1)))) [negated conjecture 1]
12. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
13. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
16. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
21. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
38. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
41. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
48. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
75. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
87. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
88. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
95. ! [X0] : (v1_xreal_0(X0) => ! [X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & ~v2_xxreal_0(X1)) => k6_numbers = k3_xcmplx_0(X1,X0))) [input l128_xreal_1]
108. ? [X0] : (? [X1] : ((k6_numbers != k3_xcmplx_0(X0,X1) & (r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,k6_numbers))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
109. ? [X0] : (? [X1] : (k6_numbers != k3_xcmplx_0(X0,X1) & r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 108]
113. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
114. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 13]
125. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 88]
126. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 125]
127. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 87]
128. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 127]
137. ! [X0] : (! [X1] : (k6_numbers = k3_xcmplx_0(X1,X0) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | v2_xxreal_0(X1))) | ~v1_xreal_0(X0)) [ennf transformation 95]
138. ! [X0] : (! [X1] : (k6_numbers = k3_xcmplx_0(X1,X0) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | v2_xxreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 137]
143. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 16]
154. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 41]
155. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 154]
156. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 38]
157. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 156]
198. ? [X0] : (? [X1] : (k6_numbers != k3_xcmplx_0(X0,X1) & r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (k6_numbers != k3_xcmplx_0(sK0,X1) & r1_xxreal_0(k6_numbers,sK0) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
199. ? [X1] : (k6_numbers != k3_xcmplx_0(sK0,X1) & r1_xxreal_0(k6_numbers,sK0) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(X1)) => (k6_numbers != k3_xcmplx_0(sK0,sK1) & r1_xxreal_0(k6_numbers,sK0) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) [choice axiom]
200. (k6_numbers != k3_xcmplx_0(sK0,sK1) & r1_xxreal_0(k6_numbers,sK0) & r1_xxreal_0(sK0,k6_numbers) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 109,199,198]
238. v1_xreal_0(sK0) [cnf transformation 200] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
239. v1_xreal_0(sK1) [cnf transformation 200] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
240. r1_xxreal_0(sK0,k6_numbers) [cnf transformation 200] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
241. r1_xxreal_0(k6_numbers,sK0) [cnf transformation 200] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
242. k6_numbers != k3_xcmplx_0(sK0,sK1) [cnf transformation 200] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
245. v1_xboole_0(k1_xboole_0) [cnf transformation 48] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
247. k6_numbers = k1_xboole_0 [cnf transformation 75] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
249. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 21] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
262. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 113] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
263. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 114] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
273. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 126] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
274. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 128] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
279. k6_numbers = k3_xcmplx_0(X1,X0) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 138] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
284. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 143] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
294. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 155] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
297. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 157] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
362. 1 <=> k6_numbers = k3_xcmplx_0(sK0,sK1) [avatar definition]
364. k6_numbers != k3_xcmplx_0(sK0,sK1) <- (~1) [avatar component clause 362] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
365. ~1 [avatar split clause 242,362]
367. 2 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
369. r1_xxreal_0(k6_numbers,sK0) <- (2) [avatar component clause 367] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
370. 2 [avatar split clause 241,367]
372. 3 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
374. r1_xxreal_0(sK0,k6_numbers) <- (3) [avatar component clause 372] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
375. 3 [avatar split clause 240,372]
377. 4 <=> v1_xreal_0(sK1) [avatar definition]
379. v1_xreal_0(sK1) <- (4) [avatar component clause 377] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
380. 4 [avatar split clause 239,377]
382. 5 <=> v1_xreal_0(sK0) [avatar definition]
384. v1_xreal_0(sK0) <- (5) [avatar component clause 382] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
385. 5 [avatar split clause 238,382]
397. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
399. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 397] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
400. 8 [avatar split clause 245,397]
407. 10 <=> k6_numbers = k1_xboole_0 [avatar definition]
409. k6_numbers = k1_xboole_0 <- (10) [avatar component clause 407] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
410. 10 [avatar split clause 247,407]
417. 12 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
419. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (12) [avatar component clause 417] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
420. 12 [avatar split clause 249,417]
638. v1_xboole_0(k6_numbers) <- (8, 10) [backward demodulation 399,409] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
640. 57 <=> v1_xboole_0(k6_numbers) [avatar definition]
642. v1_xboole_0(k6_numbers) <- (57) [avatar component clause 640] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
643. 57 | ~8 | ~10 [avatar split clause 638,407,397,640]
644. m1_subset_1(k6_numbers,k4_ordinal1) <- (10, 12) [forward demodulation 419,409] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
646. 58 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
648. m1_subset_1(k6_numbers,k4_ordinal1) <- (58) [avatar component clause 646] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
649. 58 | ~10 | ~12 [avatar split clause 644,417,407,646]
825. 86 <=> v7_ordinal1(k6_numbers) [avatar definition]
827. v7_ordinal1(k6_numbers) <- (86) [avatar component clause 825] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
893. v7_ordinal1(k6_numbers) <- (58) [resolution 284,648] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
896. 86 | ~58 [avatar split clause 893,646,825]
1003. v1_xxreal_0(k6_numbers) <- (86) [resolution 827,263] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1004. v1_xreal_0(k6_numbers) <- (86) [resolution 827,262] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1009. 111 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1012. 111 | ~86 [avatar split clause 1003,825,1009]
1014. 112 <=> v1_xreal_0(k6_numbers) [avatar definition]
1017. 112 | ~86 [avatar split clause 1004,825,1014]
1079. 120 <=> v2_xxreal_0(sK0) [avatar definition]
1083. 121 <=> v3_xxreal_0(sK0) [avatar definition]
1084. ~v3_xxreal_0(sK0) <- (~121) [avatar component clause 1083] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1243. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (2) [resolution 273,369] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1256. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 5) [subsumption resolution 1243,384] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1258. 148 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1260. v3_xxreal_0(k6_numbers) <- (148) [avatar component clause 1258] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1261. ~112 | 148 | ~121 | ~2 | ~5 [avatar split clause 1256,382,367,1083,1258,1014]
1270. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (148) [resolution 1260,294] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1273. 150 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1274. v2_xxreal_0(k6_numbers) <- (150) [avatar component clause 1273] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1277. ~v1_xxreal_0(k6_numbers) <- (57, 148) [subsumption resolution 1270,642] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1278. ~111 | ~57 | ~148 [avatar split clause 1277,1258,640,1009]
1284. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (3) [resolution 274,374] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1300. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 5) [subsumption resolution 1284,384] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1301. ~112 | 150 | ~120 | ~3 | ~5 [avatar split clause 1300,382,372,1079,1273,1014]
1450. k6_numbers != k6_numbers | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) | v2_xxreal_0(sK0) | ~v1_xreal_0(sK1) <- (~1) [superposition 364,279] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1468. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) | v2_xxreal_0(sK0) | ~v1_xreal_0(sK1) <- (~1) [trivial inequality removal 1450] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1475. v3_xxreal_0(sK0) | v2_xxreal_0(sK0) | ~v1_xreal_0(sK1) <- (~1, 5) [subsumption resolution 1468,384] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1476. v2_xxreal_0(sK0) | ~v1_xreal_0(sK1) <- (~1, 5, ~121) [subsumption resolution 1475,1084] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1500. v2_xxreal_0(sK0) <- (~1, 4, 5, ~121) [subsumption resolution 1476,379] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1501. 120 | 1 | ~4 | ~5 | 121 [avatar split clause 1500,1083,382,377,362,1079]
1533. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (150) [resolution 1274,297] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1535. ~v1_xxreal_0(k6_numbers) <- (57, 150) [subsumption resolution 1533,642] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1536. ~111 | ~57 | ~150 [avatar split clause 1535,1273,640,1009]
1537. $false [avatar sat refutation 365,370,375,380,385,400,410,420,643,649,896,1012,1017,1261,1278,1301,1501,1536] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:955,thDist:-955}
% SZS output end Proof for xreal_1__t126_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 95
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 123
% Generated clauses: 661
% Active clauses: 153
% Passive clauses: 254
% Final active clauses: 141
% Final passive clauses: 81
% Main loop iterations started: 167
% 
% >>> Simplifying Inferences
% Duplicate literals: 69
% Trivial inequalities: 1
% Fw subsumption resolutions: 67
% Fw demodulations: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 31
% Forward subsumptions: 80
% 
% >>> Generating Inferences
% Binary resolution: 167
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 46
% Backward superposition: 25
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 40
% Split components: 97
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 186
% SAT solver unit clauses: 57
% SAT solver binary clauses: 72
% 
% Memory used [KB]: 5756
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% congruence closure: 0.002 s
% forward subsumption: 0.001 s
% resolution: 0.001 s
% superposition: 0.001 s
% other: 0.007 s ( own 0.001 s ) 

