i: [1,48,0,2,1,1,0,0] t8_xreal_1
a: [3,82,0,2,1,48]
s: 82
i: [1,47,0,2,1,1,0,0] t8_xreal_1
a: [3,87,0,2,1,47]
s: 87
i: [1,46,0,2,1,1,0,0] t8_xreal_1
a: [3,92,0,2,1,46]
s: 92
i: [1,45,0,2,1,1,0,0] t8_xreal_1
a: [3,97,0,2,1,45]
s: 97
i: [1,50,0,3,1,1,0,0] t8_xreal_1
a: [3,72,0,3,1,50]
s: 72
i: [1,49,0,3,1,1,0,0] t8_xreal_1
a: [3,77,0,3,1,49]
s: 77
i: [1,51,0,7,1,1,0,0] t8_xreal_1
a: [3,67,0,7,1,51]
s: 67
i: [1,61,0,2,1,0,0,2] rc1_xxreal_0
a: [3,109,0,2,1,61]
s: 109
i: [1,62,0,2,1,0,0,255] rc1_xcmplx_0
a: [3,114,0,2,1,62]
s: 114
i: [1,63,0,2,1,0,0,1] rc1_xreal_0
a: [3,119,0,2,1,63]
s: 119
i: [1,52,0,4,2,0,0,2] cc4_xreal_0
s: 52
d: [2,121,1,2,1,1,40,52,97]
a: [3,129,1,2,1,121]
s: 129
i: [1,53,0,4,2,0,0,255] cc3_xreal_0
s: 53
d: [2,122,1,2,1,1,40,52,92]
a: [3,134,1,2,1,122]
s: 134
i: [1,55,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,104,0,5,2,55]
s: 104
i: [1,57,0,8,3,0,0,255] fc2_xcmplx_0
s: 57
d: [2,123,1,2,1,1,40,52,87]
a: [3,139,1,2,1,123]
s: 139
i: [1,59,0,8,3,0,0,1] fc5_xreal_0
s: 59
d: [2,124,1,2,1,1,40,52,82]
a: [3,144,1,2,1,124]
s: 144
i: [1,56,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 56
i: [1,58,0,11,3,0,0,255] commutativity_k2_xcmplx_0
s: 58
d: [2,151,1,2,1,1,40,53,97]
a: [3,159,1,2,1,151]
s: 159
i: [1,60,0,17,4,0,0,255] spc6_arithm
s: 60
d: [2,152,1,2,1,1,40,53,92]
a: [3,164,1,2,1,152]
s: 164
i: [1,54,0,21,7,0,0,1] l8_xreal_1
s: 54
d: [2,223,1,0,0,7,52,97,92,87,82,77,72,67,54]
e: 223
d: [2,224,1,14,6,1,40,54,67]
d: [2,235,1,11,5,2,21,224,72]
d: [2,236,1,8,4,3,21,235,77]
d: [2,237,1,6,3,4,21,236,82]
d: [2,238,1,4,2,5,21,237,87]
d: [2,239,1,2,1,6,21,238,92]
d: [2,240,1,0,0,7,21,239,97]
e: 240
f: 240,45,46,47,48,49,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t8_xreal_1
% SZS output start Proof for xreal_1__t8_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ! [X2] : (v1_xreal_0(X2) => ! [X3] : (v1_xreal_0(X3) => ~(r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,X0)))))) [input t8_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ! [X2] : (v1_xreal_0(X2) => ! [X3] : (v1_xreal_0(X3) => ~(r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,X0)))))) [negated conjecture 1]
15. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ! [X2] : (v1_xreal_0(X2) => ! [X3] : (v1_xreal_0(X3) => ~(r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,X0)))))) [input l8_xreal_1]
16. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,X0)) & v1_xreal_0(X3)) & v1_xreal_0(X2)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
17. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,X0) & v1_xreal_0(X3)) & v1_xreal_0(X2)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 16]
20. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((~r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X0,X2)) | ~r1_xxreal_0(X2,X3) | r1_xxreal_0(X1,X0)) | ~v1_xreal_0(X3)) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 15]
21. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (~r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X0,X2)) | ~r1_xxreal_0(X2,X3) | r1_xxreal_0(X1,X0) | ~v1_xreal_0(X3)) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 20]
34. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,X0) & v1_xreal_0(X3)) & v1_xreal_0(X2)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(sK0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,sK0) & v1_xreal_0(X3)) & v1_xreal_0(X2)) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
35. ? [X1] : (? [X2] : (? [X3] : (r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(sK0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,sK0) & v1_xreal_0(X3)) & v1_xreal_0(X2)) & v1_xreal_0(X1)) => (? [X2] : (? [X3] : (r1_xxreal_0(k2_xcmplx_0(sK1,X3),k2_xcmplx_0(sK0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(sK1,sK0) & v1_xreal_0(X3)) & v1_xreal_0(X2)) & v1_xreal_0(sK1)) [choice axiom]
36. ? [X2] : (? [X3] : (r1_xxreal_0(k2_xcmplx_0(sK1,X3),k2_xcmplx_0(sK0,X2)) & r1_xxreal_0(X2,X3) & ~r1_xxreal_0(sK1,sK0) & v1_xreal_0(X3)) & v1_xreal_0(X2)) => (? [X3] : (r1_xxreal_0(k2_xcmplx_0(sK1,X3),k2_xcmplx_0(sK0,sK2)) & r1_xxreal_0(sK2,X3) & ~r1_xxreal_0(sK1,sK0) & v1_xreal_0(X3)) & v1_xreal_0(sK2)) [choice axiom]
37. ? [X3] : (r1_xxreal_0(k2_xcmplx_0(sK1,X3),k2_xcmplx_0(sK0,sK2)) & r1_xxreal_0(sK2,X3) & ~r1_xxreal_0(sK1,sK0) & v1_xreal_0(X3)) => (r1_xxreal_0(k2_xcmplx_0(sK1,sK3),k2_xcmplx_0(sK0,sK2)) & r1_xxreal_0(sK2,sK3) & ~r1_xxreal_0(sK1,sK0) & v1_xreal_0(sK3)) [choice axiom]
38. (((r1_xxreal_0(k2_xcmplx_0(sK1,sK3),k2_xcmplx_0(sK0,sK2)) & r1_xxreal_0(sK2,sK3) & ~r1_xxreal_0(sK1,sK0) & v1_xreal_0(sK3)) & v1_xreal_0(sK2)) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 17,37,36,35,34]
45. v1_xreal_0(sK0) [cnf transformation 38] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
46. v1_xreal_0(sK1) [cnf transformation 38] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
47. v1_xreal_0(sK2) [cnf transformation 38] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
48. v1_xreal_0(sK3) [cnf transformation 38] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
49. ~r1_xxreal_0(sK1,sK0) [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
50. r1_xxreal_0(sK2,sK3) [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
51. r1_xxreal_0(k2_xcmplx_0(sK1,sK3),k2_xcmplx_0(sK0,sK2)) [cnf transformation 38] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
54. ~r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X0,X2)) | ~r1_xxreal_0(X2,X3) | r1_xxreal_0(X1,X0) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 21] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. 1 <=> r1_xxreal_0(k2_xcmplx_0(sK1,sK3),k2_xcmplx_0(sK0,sK2)) [avatar definition]
67. r1_xxreal_0(k2_xcmplx_0(sK1,sK3),k2_xcmplx_0(sK0,sK2)) <- (1) [avatar component clause 65] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
68. 1 [avatar split clause 51,65]
70. 2 <=> r1_xxreal_0(sK2,sK3) [avatar definition]
72. r1_xxreal_0(sK2,sK3) <- (2) [avatar component clause 70] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
73. 2 [avatar split clause 50,70]
75. 3 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
77. ~r1_xxreal_0(sK1,sK0) <- (~3) [avatar component clause 75] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
78. ~3 [avatar split clause 49,75]
80. 4 <=> v1_xreal_0(sK3) [avatar definition]
82. v1_xreal_0(sK3) <- (4) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
83. 4 [avatar split clause 48,80]
85. 5 <=> v1_xreal_0(sK2) [avatar definition]
87. v1_xreal_0(sK2) <- (5) [avatar component clause 85] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
88. 5 [avatar split clause 47,85]
90. 6 <=> v1_xreal_0(sK1) [avatar definition]
92. v1_xreal_0(sK1) <- (6) [avatar component clause 90] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
93. 6 [avatar split clause 46,90]
95. 7 <=> v1_xreal_0(sK0) [avatar definition]
97. v1_xreal_0(sK0) <- (7) [avatar component clause 95] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
98. 7 [avatar split clause 45,95]
224. ~r1_xxreal_0(sK2,sK3) | r1_xxreal_0(sK1,sK0) | ~v1_xreal_0(sK3) | ~v1_xreal_0(sK2) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (1) [resolution 54,67] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
235. r1_xxreal_0(sK1,sK0) | ~v1_xreal_0(sK3) | ~v1_xreal_0(sK2) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (1, 2) [subsumption resolution 224,72] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
236. ~v1_xreal_0(sK3) | ~v1_xreal_0(sK2) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (1, 2, ~3) [subsumption resolution 235,77] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
237. ~v1_xreal_0(sK2) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (1, 2, ~3, 4) [subsumption resolution 236,82] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
238. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (1, 2, ~3, 4, 5) [subsumption resolution 237,87] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
239. ~v1_xreal_0(sK0) <- (1, 2, ~3, 4, 5, 6) [subsumption resolution 238,92] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
240. $false <- (1, 2, ~3, 4, 5, 6, 7) [subsumption resolution 239,97] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
241. ~1 | ~2 | 3 | ~4 | ~5 | ~6 | ~7 [avatar contradiction clause 240]
252. $false [avatar sat refutation 68,73,78,83,88,93,98,241] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:180,thDist:-180}
% SZS output end Proof for xreal_1__t8_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 115
% Active clauses: 25
% Passive clauses: 41
% Final active clauses: 25
% Final passive clauses: 16
% Main loop iterations started: 25
% 
% >>> Simplifying Inferences
% Duplicate literals: 17
% Fw subsumption resolutions: 25
% 
% >>> Deletion Inferences
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 16
% Unit resulting resolution: 1
% Factoring: 1
% Forward superposition: 10
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 10
% SAT solver binary clauses: 11
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

