i: [1,51,0,15,3,1,0,0] t28_zfmisc_1
a: [3,82,0,3,1,51]
s: 82
a: [3,86,0,3,1,51]
s: 86
i: [1,52,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,94,0,2,1,52]
s: 94
i: [1,53,0,2,1,0,0,255] fc1_xboole_0
a: [3,99,0,2,1,53]
s: 99
i: [1,71,0,2,1,0,0,2] rc2_xboole_0
a: [3,109,0,2,1,71]
s: 109
i: [1,72,0,2,1,0,0,2] rc1_xboole_0
a: [3,114,0,2,1,72]
s: 114
a: [3,79,0,9,1,51]
s: 79
d: [2,116,0,9,1,2,22,79,86]
d: [2,117,0,9,1,3,22,116,82]
a: [3,121,0,9,1,117]
s: 121
i: [1,55,0,3,1,0,0,1] fc2_xboole_0
s: 55
i: [1,54,0,3,1,0,0,255] d2_xboole_0
a: [3,104,0,3,1,54]
s: 104
r: 94
i: [1,67,0,4,1,0,0,1] fc1_zfmisc_1
s: 67
i: [1,68,0,4,1,0,0,255] fc3_xboole_0
s: 68
i: [1,74,0,4,1,0,0,1] d1_tarski
s: 74
i: [1,56,0,5,2,0,0,255] t6_boole
s: 56
i: [1,63,0,5,2,0,0,1] t7_boole
s: 63
i: [1,70,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 70
d: [2,127,1,4,1,0,40,70,74]
s: 127
i: [1,62,0,7,3,0,0,1] t8_boole
s: 62
i: [1,69,0,7,1,0,0,255] commutativity_k2_tarski
s: 69
i: [1,75,0,7,2,0,0,1] d1_tarski
s: 75
i: [1,57,0,10,1,0,0,255] d5_tarski
s: 57
d: [2,124,1,7,3,0,44,56,56]
s: 124
i: [1,64,0,10,2,0,0,1] l46_zfmisc_1
s: 64
i: [1,65,0,10,2,0,0,1] l46_zfmisc_1
s: 65
i: [1,66,0,13,3,0,0,1] l46_zfmisc_1
s: 66
d: [2,139,1,0,0,1,52,74,74,121,66]
e: 139
d: [2,140,1,8,2,1,40,66,121]
d: [2,146,1,4,1,1,21,140,74]
d: [2,147,1,0,0,1,21,146,74]
e: 147
r: 86
r: 121
a: [3,87,0,3,1,51]
s: 87
a: [3,120,0,9,1,117]
s: 120
a: [3,78,0,9,1,51]
d: [2,149,0,9,1,2,22,78,82]
a: [3,153,0,9,1,149]
s: 153
d: [2,169,1,4,1,1,40,153,65]
a: [3,176,1,4,1,169]
s: 176
d: [2,188,2,0,0,2,52,87,176,75]
e: 188
d: [2,189,2,3,1,1,40,176,75]
d: [2,193,2,0,0,2,21,189,87]
e: 193
r: 82
r: 153
r: 176
a: [3,83,0,3,1,51]
s: 83
s: 78
r: 87
s: 86
r: 78
s: 153
s: 176
d: [2,216,0,9,1,2,23,78,86]
a: [3,220,0,9,1,216]
s: 220
i: [1,60,0,14,3,0,0,1] d1_tarski
s: 60
d: [2,196,1,4,1,1,40,78,64]
a: [3,203,1,4,1,196]
s: 203
d: [2,242,2,0,0,2,52,83,203,75]
e: 242
d: [2,243,2,3,1,1,40,203,75]
d: [2,247,2,0,0,2,21,243,83]
e: 247
i: [1,50,0,12,2,1,0,0] t28_zfmisc_1
i: [1,49,0,12,2,1,0,0] t28_zfmisc_1
d: [2,195,1,4,1,1,40,78,65]
f: 247,196,195,193,147,117,49,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t28_zfmisc_1
% SZS output start Proof for zfmisc_1__t28_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3))) <=> (X1 = X3 & X0 = X2)) [input t28_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3))) <=> (X1 = X3 & X0 = X2)) [negated conjecture 1]
5. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
18. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [input l46_zfmisc_1]
24. ~! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3))) <=> (X1 = X3 & X0 = X2)) [flattening 2]
26. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 5]
29. ! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [flattening 18]
30. ? [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3))) <~> (X1 = X3 & X0 = X2)) [ennf transformation 24]
35. ? [X0,X1,X2,X3] : (((X1 != X3 | X0 != X2) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3)))) & ((X1 = X3 & X0 = X2) | r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3))))) [nnf transformation 30]
36. ? [X0,X1,X2,X3] : ((X1 != X3 | X0 != X2 | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3)))) & ((X1 = X3 & X0 = X2) | r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3))))) [flattening 35]
37. ? [X0,X1,X2,X3] : ((X1 != X3 | X0 != X2 | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3)))) & ((X1 = X3 & X0 = X2) | r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X3))))) => ((sK1 != sK3 | sK0 != sK2 | ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3)))) & ((sK1 = sK3 & sK0 = sK2) | r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3))))) [choice axiom]
38. (sK1 != sK3 | sK0 != sK2 | ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3)))) & ((sK1 = sK3 & sK0 = sK2) | r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3)))) [skolemisation 36,37]
39. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 26]
40. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 39]
41. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 40,41]
43. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | (~r2_hidden(X1,X3) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [nnf transformation 29]
44. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [flattening 43]
49. sK0 = sK2 | r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3))) [cnf transformation 38] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. sK1 = sK3 | r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3))) [cnf transformation 38] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. sK1 != sK3 | sK0 != sK2 | ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3))) [cnf transformation 38] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 42] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
59. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 42] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
64. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X0,X2) [cnf transformation 44] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X1,X3) [cnf transformation 44] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
66. r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2) [cnf transformation 44] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 59] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:2,thDist:-2}
74. r2_hidden(X3,k1_tarski(X3)) [equality resolution 73] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
75. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 58] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
77. 1 <=> r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3))) [avatar definition]
78. r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3))) <- (1) [avatar component clause 77] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK3))) <- (~1) [avatar component clause 77] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
81. 2 <=> sK0 = sK2 [avatar definition]
82. sK0 = sK2 <- (2) [avatar component clause 81] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. sK0 != sK2 <- (~2) [avatar component clause 81] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
85. 3 <=> sK1 = sK3 [avatar definition]
86. sK1 = sK3 <- (3) [avatar component clause 85] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. sK1 != sK3 <- (~3) [avatar component clause 85] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
88. ~1 | ~2 | ~3 [avatar split clause 51,85,81,77]
89. 1 | 3 [avatar split clause 50,85,77]
90. 1 | 2 [avatar split clause 49,81,77]
116. ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK2),k1_tarski(sK1))) <- (~1, 3) [forward demodulation 79,86] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
117. ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK0),k1_tarski(sK1))) <- (~1, 2, 3) [forward demodulation 116,82] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
119. 9 <=> r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK0),k1_tarski(sK1))) [avatar definition]
121. ~r2_hidden(k4_tarski(sK0,sK1),k2_zfmisc_1(k1_tarski(sK0),k1_tarski(sK1))) <- (~9) [avatar component clause 119] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
122. ~9 | 1 | ~2 | ~3 [avatar split clause 117,85,81,77,119]
140. ~r2_hidden(sK1,k1_tarski(sK1)) | ~r2_hidden(sK0,k1_tarski(sK0)) <- (~9) [resolution 66,121] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
146. ~r2_hidden(sK0,k1_tarski(sK0)) <- (~9) [subsumption resolution 140,74] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
147. $false <- (~9) [subsumption resolution 146,74] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
148. 9 [avatar contradiction clause 147]
174. 13 <=> r2_hidden(sK1,k1_tarski(sK3)) [avatar definition]
176. r2_hidden(sK1,k1_tarski(sK3)) <- (13) [avatar component clause 174] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
189. sK1 = sK3 <- (13) [resolution 176,75] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
193. $false <- (~3, 13) [subsumption resolution 189,87] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
194. 3 | ~13 [avatar contradiction clause 193]
195. r2_hidden(sK1,k1_tarski(sK3)) <- (1) [resolution 78,65] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
196. r2_hidden(sK0,k1_tarski(sK2)) <- (1) [resolution 78,64] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
199. 13 | ~1 [avatar split clause 195,77,174]
201. 16 <=> r2_hidden(sK0,k1_tarski(sK2)) [avatar definition]
203. r2_hidden(sK0,k1_tarski(sK2)) <- (16) [avatar component clause 201] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
204. 16 | ~1 [avatar split clause 196,77,201]
243. sK0 = sK2 <- (16) [resolution 203,75] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
247. $false <- (~2, 16) [subsumption resolution 243,83] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
248. 2 | ~16 [avatar contradiction clause 247]
249. $false [avatar sat refutation 88,89,90,122,148,194,199,204,248] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:136,thDist:-136}
% SZS output end Proof for zfmisc_1__t28_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 115
% Active clauses: 33
% Passive clauses: 59
% Final active clauses: 25
% Final passive clauses: 21
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 4
% Fw demodulations: 5
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 14
% 
% >>> Generating Inferences
% Binary resolution: 33
% Unit resulting resolution: 3
% Forward superposition: 4
% Backward superposition: 6
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 32
% SAT solver unit clauses: 7
% SAT solver binary clauses: 18
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

