i: [1,75,0,15,4,1,0,0] t31_mcart_1
a: [3,110,0,3,1,75]
s: 110
a: [3,117,0,6,1,75]
s: 117
d: [2,163,0,6,1,2,22,117,110]
a: [3,167,0,6,1,163]
s: 167
i: [1,76,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,126,0,2,1,76]
s: 126
i: [1,77,0,2,1,0,0,1] fc1_xboole_0
a: [3,131,0,2,1,77]
s: 131
i: [1,96,0,2,1,0,0,2] rc2_xboole_0
a: [3,141,0,2,1,96]
s: 141
i: [1,97,0,2,1,0,0,2] rc1_xboole_0
a: [3,146,0,2,1,97]
s: 146
i: [1,99,0,2,1,0,0,3] rc1_relat_1
a: [3,151,0,2,1,99]
s: 151
i: [1,98,0,2,1,0,0,3] rc1_relat_1
a: [3,156,0,2,1,98]
s: 156
i: [1,100,0,2,1,0,0,255] rc2_relat_1
a: [3,161,0,2,1,100]
s: 161
i: [1,78,0,3,1,0,0,255] d2_xboole_0
a: [3,136,0,3,1,78]
s: 136
r: 126
i: [1,79,0,4,2,0,0,3] cc1_relat_1
s: 79
i: [1,91,0,4,1,0,0,255] existence_m1_subset_1
s: 91
d: [2,170,1,2,1,1,40,79,146]
a: [3,175,1,2,1,170]
s: 175
i: [1,92,0,4,1,0,0,2] fc6_relat_1
s: 92
i: [1,80,0,5,2,0,0,1] t6_boole
s: 80
d: [2,184,1,8,2,1,44,167,80]
s: 184
i: [1,89,0,5,2,0,0,255] t7_boole
s: 89
d: [2,171,1,2,1,1,40,79,131]
a: [3,180,1,2,1,171]
s: 180
i: [1,101,0,5,1,0,0,2] t90_zfmisc_1
s: 101
i: [1,102,0,5,1,0,0,2] t90_zfmisc_1
s: 102
i: [1,83,0,6,2,0,0,255] t1_subset
s: 83
i: [1,93,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 93
i: [1,88,0,7,3,0,0,1] t8_boole
s: 88
i: [1,84,0,8,3,0,0,255] t2_subset
s: 84
i: [1,94,0,8,3,0,0,2] fc10_subset_1
s: 94
d: [2,190,1,6,2,0,40,84,91]
s: 190
i: [1,90,0,10,1,0,0,1] d3_zfmisc_1
s: 90
d: [2,201,1,5,1,0,44,92,90]
s: 201
i: [1,85,0,11,3,0,0,2] t90_zfmisc_1
s: 85
i: [1,95,0,11,4,0,0,1] fc11_subset_1
s: 95
d: [2,199,1,6,1,0,44,90,101]
s: 199
d: [2,202,1,6,1,0,44,101,90]
s: 202
d: [2,196,1,8,1,0,44,90,101]
d: [2,203,1,6,1,0,22,196,102]
s: 203
d: [2,214,2,3,1,1,44,167,203]
d: [2,219,2,0,0,1,19,214]
e: 219
d: [2,211,2,0,0,1,52,167,203]
e: 211
d: [2,212,2,0,0,2,52,131,203,184]
e: 212
d: [2,215,2,5,2,1,44,184,203]
d: [2,218,2,2,1,1,19,215]
d: [2,223,2,0,0,2,21,218,131]
e: 223
r: 110
r: 184
r: 167
a: [3,109,0,3,1,75]
s: 109
a: [3,114,0,3,1,75]
s: 114
s: 117
d: [2,225,0,6,1,2,22,117,114]
d: [2,226,0,0,0,2,21,225,199]
e: 226
r: 114
a: [3,106,0,3,1,75]
s: 106
s: 117
d: [2,228,0,6,1,2,22,117,106]
a: [3,232,0,6,1,228]
s: 232
d: [2,234,1,8,2,1,44,232,80]
s: 234
d: [2,197,1,8,1,0,44,90,102]
d: [2,204,1,6,1,0,22,197,102]
s: 204
d: [2,238,2,3,1,1,44,232,204]
d: [2,243,2,0,0,1,19,238]
e: 243
d: [2,235,2,0,0,1,52,232,204]
e: 235
d: [2,236,2,0,0,2,52,131,204,234]
e: 236
d: [2,241,2,5,2,1,44,234,204]
d: [2,242,2,2,1,1,19,241]
d: [2,247,2,0,0,2,21,242,131]
e: 247
r: 106
r: 234
r: 232
a: [3,105,0,3,1,75]
s: 105
a: [3,113,0,3,1,75]
s: 113
a: [3,118,0,6,1,75]
s: 118
d: [2,249,1,8,4,1,44,95,118]
d: [2,251,1,6,3,2,21,249,131]
a: [3,259,1,2,1,251]
s: 259
d: [2,265,2,0,0,3,52,131,109,259,88]
e: 265
d: [2,266,2,0,0,3,52,131,109,259,88]
e: 266
d: [2,267,2,0,0,2,52,109,259,80]
e: 267
r: 259
a: [3,258,1,2,1,251]
s: 258
a: [3,263,1,2,1,251]
s: 263
d: [2,272,2,0,0,3,52,131,113,263,88]
e: 272
d: [2,273,2,0,0,3,52,131,113,263,88]
e: 273
d: [2,274,2,0,0,2,52,113,263,80]
e: 274
r: 263
a: [3,255,1,2,1,251]
s: 255
d: [2,279,2,0,0,3,52,131,105,255,88]
e: 279
d: [2,280,2,0,0,3,52,131,105,255,88]
e: 280
d: [2,281,2,0,0,2,52,105,255,80]
e: 281
i: [1,74,0,9,2,1,0,0] t31_mcart_1
i: [1,73,0,9,2,1,0,0] t31_mcart_1
i: [1,72,0,9,2,1,0,0] t31_mcart_1
f: 281,274,267,251,235,228,226,211,163,77,72,73,74,75
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t31_mcart_1
% SZS output start Proof for mcart_1__t31_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : ((k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) <=> k1_xboole_0 != k3_zfmisc_1(X0,X1,X2)) [input t31_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : ((k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) <=> k1_xboole_0 != k3_zfmisc_1(X0,X1,X2)) [negated conjecture 1]
8. ! [X0] : ! [X1] : ! [X2] : k3_zfmisc_1(X0,X1,X2) = k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2) [input d3_zfmisc_1]
16. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k3_zfmisc_1(X0,X1,X2))) [input fc11_subset_1]
17. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
25. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
28. ! [X0] : ! [X1] : (k1_xboole_0 = k2_zfmisc_1(X0,X1) <=> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t90_zfmisc_1]
29. ~! [X0,X1,X2] : ((k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) <=> k1_xboole_0 != k3_zfmisc_1(X0,X1,X2)) [flattening 2]
32. ! [X0,X1] : (k1_xboole_0 = k2_zfmisc_1(X0,X1) <=> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 28]
35. ! [X0,X1,X2] : k3_zfmisc_1(X0,X1,X2) = k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2) [flattening 8]
39. ? [X0,X1,X2] : ((k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0) <~> k1_xboole_0 != k3_zfmisc_1(X0,X1,X2)) [ennf transformation 29]
41. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 25]
54. ! [X0,X1,X2] : (~v1_xboole_0(k3_zfmisc_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 16]
55. ! [X0,X1,X2] : (~v1_xboole_0(k3_zfmisc_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 54]
56. ? [X0,X1,X2] : ((k1_xboole_0 = k3_zfmisc_1(X0,X1,X2) | (k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) & (k1_xboole_0 != k3_zfmisc_1(X0,X1,X2) | (k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0))) [nnf transformation 39]
57. ? [X0,X1,X2] : ((k1_xboole_0 = k3_zfmisc_1(X0,X1,X2) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0) & (k1_xboole_0 != k3_zfmisc_1(X0,X1,X2) | (k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0))) [flattening 56]
58. ? [X0,X1,X2] : ((k1_xboole_0 = k3_zfmisc_1(X0,X1,X2) | k1_xboole_0 = X2 | k1_xboole_0 = X1 | k1_xboole_0 = X0) & (k1_xboole_0 != k3_zfmisc_1(X0,X1,X2) | (k1_xboole_0 != X2 & k1_xboole_0 != X1 & k1_xboole_0 != X0))) => ((k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) | k1_xboole_0 = sK2 | k1_xboole_0 = sK1 | k1_xboole_0 = sK0) & (k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) | (k1_xboole_0 != sK2 & k1_xboole_0 != sK1 & k1_xboole_0 != sK0))) [choice axiom]
59. (k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) | k1_xboole_0 = sK2 | k1_xboole_0 = sK1 | k1_xboole_0 = sK0) & (k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) | (k1_xboole_0 != sK2 & k1_xboole_0 != sK1 & k1_xboole_0 != sK0)) [skolemisation 57,58]
60. ! [X0,X1] : ((k1_xboole_0 = k2_zfmisc_1(X0,X1) | (k1_xboole_0 != X1 & k1_xboole_0 != X0)) & ((k1_xboole_0 = X1 | k1_xboole_0 = X0) | k1_xboole_0 != k2_zfmisc_1(X0,X1))) [nnf transformation 32]
61. ! [X0,X1] : ((k1_xboole_0 = k2_zfmisc_1(X0,X1) | (k1_xboole_0 != X1 & k1_xboole_0 != X0)) & (k1_xboole_0 = X1 | k1_xboole_0 = X0 | k1_xboole_0 != k2_zfmisc_1(X0,X1))) [flattening 60]
72. k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) | k1_xboole_0 != sK0 [cnf transformation 59] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
73. k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) | k1_xboole_0 != sK1 [cnf transformation 59] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
74. k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) | k1_xboole_0 != sK2 [cnf transformation 59] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
75. k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) | k1_xboole_0 = sK2 | k1_xboole_0 = sK1 | k1_xboole_0 = sK0 [cnf transformation 59] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
77. v1_xboole_0(k1_xboole_0) [cnf transformation 17] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
80. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 41] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
86. k1_xboole_0 = k2_zfmisc_1(X0,X1) | k1_xboole_0 != X0 [cnf transformation 61] {a:0,w:8,wCS:3200,sine:2,thAx:0,allAx:1,thDist:-1}
87. k1_xboole_0 = k2_zfmisc_1(X0,X1) | k1_xboole_0 != X1 [cnf transformation 61] {a:0,w:8,wCS:3200,sine:2,thAx:0,allAx:1,thDist:-1}
90. k3_zfmisc_1(X0,X1,X2) = k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2) [cnf transformation 35] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
95. ~v1_xboole_0(k3_zfmisc_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 55] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. k1_xboole_0 = k2_zfmisc_1(X0,k1_xboole_0) [equality resolution 87] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
102. k1_xboole_0 = k2_zfmisc_1(k1_xboole_0,X1) [equality resolution 86] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
104. 1 <=> k1_xboole_0 = sK0 [avatar definition]
105. k1_xboole_0 != sK0 <- (~1) [avatar component clause 104] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
106. k1_xboole_0 = sK0 <- (1) [avatar component clause 104] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
108. 2 <=> k1_xboole_0 = sK1 [avatar definition]
109. k1_xboole_0 != sK1 <- (~2) [avatar component clause 108] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
110. k1_xboole_0 = sK1 <- (2) [avatar component clause 108] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
112. 3 <=> k1_xboole_0 = sK2 [avatar definition]
113. k1_xboole_0 != sK2 <- (~3) [avatar component clause 112] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
114. k1_xboole_0 = sK2 <- (3) [avatar component clause 112] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. 4 <=> k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) [avatar definition]
117. k1_xboole_0 != k3_zfmisc_1(sK0,sK1,sK2) <- (~4) [avatar component clause 116] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
118. k1_xboole_0 = k3_zfmisc_1(sK0,sK1,sK2) <- (4) [avatar component clause 116] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
119. 1 | 2 | 3 | 4 [avatar split clause 75,116,112,108,104]
120. ~3 | ~4 [avatar split clause 74,116,112]
121. ~2 | ~4 [avatar split clause 73,116,108]
122. ~1 | ~4 [avatar split clause 72,116,104]
129. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
131. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 129] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
132. 6 [avatar split clause 77,129]
163. k1_xboole_0 != k3_zfmisc_1(sK0,k1_xboole_0,sK2) <- (2, ~4) [forward demodulation 117,110] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
165. 13 <=> k1_xboole_0 = k3_zfmisc_1(sK0,k1_xboole_0,sK2) [avatar definition]
167. k1_xboole_0 != k3_zfmisc_1(sK0,k1_xboole_0,sK2) <- (~13) [avatar component clause 165] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
168. ~13 | ~2 | 4 [avatar split clause 163,116,108,165]
196. k2_zfmisc_1(k1_xboole_0,X1) = k3_zfmisc_1(X0,k1_xboole_0,X1) [superposition 90,101] {a:1,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
197. k3_zfmisc_1(k1_xboole_0,X2,X3) = k2_zfmisc_1(k1_xboole_0,X3) [superposition 90,102] {a:1,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
199. k1_xboole_0 = k3_zfmisc_1(X3,X4,k1_xboole_0) [superposition 90,101] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
203. k1_xboole_0 = k3_zfmisc_1(X0,k1_xboole_0,X1) [forward demodulation 196,102] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
204. k1_xboole_0 = k3_zfmisc_1(k1_xboole_0,X2,X3) [forward demodulation 197,102] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
211. $false <- (~13) [unit resulting resolution 167,203] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
221. 13 [avatar contradiction clause 211]
225. k1_xboole_0 != k3_zfmisc_1(sK0,sK1,k1_xboole_0) <- (3, ~4) [forward demodulation 117,114] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
226. $false <- (3, ~4) [subsumption resolution 225,199] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
227. ~3 | 4 [avatar contradiction clause 226]
228. k1_xboole_0 != k3_zfmisc_1(k1_xboole_0,sK1,sK2) <- (1, ~4) [forward demodulation 117,106] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
230. 16 <=> k1_xboole_0 = k3_zfmisc_1(k1_xboole_0,sK1,sK2) [avatar definition]
232. k1_xboole_0 != k3_zfmisc_1(k1_xboole_0,sK1,sK2) <- (~16) [avatar component clause 230] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
233. ~16 | ~1 | 4 [avatar split clause 228,116,104,230]
235. $false <- (~16) [unit resulting resolution 232,204] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
245. 16 [avatar contradiction clause 235]
249. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (4) [superposition 95,118] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
251. v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (4, 6) [subsumption resolution 249,131] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
253. 17 <=> v1_xboole_0(sK0) [avatar definition]
255. v1_xboole_0(sK0) <- (17) [avatar component clause 253] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
257. 18 <=> v1_xboole_0(sK1) [avatar definition]
259. v1_xboole_0(sK1) <- (18) [avatar component clause 257] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
261. 19 <=> v1_xboole_0(sK2) [avatar definition]
263. v1_xboole_0(sK2) <- (19) [avatar component clause 261] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
264. 17 | 18 | 19 | ~4 | ~6 [avatar split clause 251,129,116,261,257,253]
267. $false <- (~2, 18) [unit resulting resolution 109,259,80] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
271. 2 | ~18 [avatar contradiction clause 267]
274. $false <- (~3, 19) [unit resulting resolution 113,263,80] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
278. 3 | ~19 [avatar contradiction clause 274]
281. $false <- (~1, 17) [unit resulting resolution 105,255,80] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
285. 1 | ~17 [avatar contradiction clause 281]
286. $false [avatar sat refutation 119,120,121,122,132,168,221,227,233,245,264,271,278,285] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:166,thDist:-166}
% SZS output end Proof for mcart_1__t31_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 27
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 130
% Active clauses: 45
% Passive clauses: 60
% Final active clauses: 35
% Final passive clauses: 11
% Main loop iterations started: 49
% 
% >>> Simplifying Inferences
% Trivial inequalities: 4
% Fw subsumption resolutions: 4
% Fw demodulations: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 13
% Forward superposition: 13
% Backward superposition: 19
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 13
% Sat splitting refutations: 18
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 37
% SAT solver unit clauses: 12
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

