i: [1,199,0,2,1,1,0,0] t4_connsp_3
a: [3,298,0,2,1,199]
s: 298
i: [1,198,0,2,1,1,0,0] t4_connsp_3
a: [3,303,0,2,1,198]
s: 303
i: [1,197,0,2,1,1,0,0] t4_connsp_3
a: [3,308,0,2,1,197]
s: 308
i: [1,201,0,3,1,1,0,0] t4_connsp_3
a: [3,288,0,3,1,201]
s: 288
i: [1,202,0,5,1,1,0,0] t4_connsp_3
a: [3,283,0,5,1,202]
s: 283
i: [1,200,0,5,1,1,0,0] t4_connsp_3
a: [3,293,0,5,1,200]
s: 293
i: [1,203,0,2,1,0,0,255] spc1_boole
a: [3,313,0,2,1,203]
s: 313
i: [1,204,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,318,0,2,1,204]
s: 318
i: [1,278,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,328,0,2,1,278]
s: 328
i: [1,279,0,2,1,0,0,2] existence_l1_struct_0
a: [3,333,0,2,1,279]
s: 333
i: [1,206,0,3,1,0,0,2] fc1_subset_1
s: 206
i: [1,272,0,3,1,0,0,2] rc2_subset_1
s: 272
i: [1,275,0,3,1,0,0,2] reflexivity_r1_tarski
s: 275
i: [1,205,0,3,1,0,0,255] d2_xboole_0
a: [3,323,0,3,1,205]
s: 323
r: 318
d: [2,335,0,2,1,2,23,318,323]
a: [3,339,0,2,1,335]
s: 339
i: [1,218,0,4,2,0,0,1] dt_l1_pre_topc
s: 218
d: [2,341,1,2,1,1,40,218,298]
a: [3,346,1,2,1,341]
s: 346
i: [1,270,0,4,1,0,0,1] existence_m1_subset_1
s: 270
d: [2,342,1,2,1,1,40,218,328]
a: [3,351,1,2,1,342]
s: 351
i: [1,274,0,4,1,0,0,255] rc3_subset_1
s: 274
i: [1,208,0,5,2,0,0,255] rc6_subset_1
s: 208
i: [1,212,0,5,2,0,0,2] rc1_subset_1
s: 212
i: [1,214,0,5,2,0,0,255] rc5_subset_1
s: 214
i: [1,215,0,5,2,0,0,255] rc5_subset_1
s: 215
i: [1,221,0,5,2,0,0,1] rc1_connsp_1
s: 221
i: [1,222,0,5,2,0,0,2] fc3_struct_0
s: 222
i: [1,240,0,5,2,0,0,1] t6_boole
s: 240
i: [1,267,0,5,2,0,0,255] t7_boole
s: 267
i: [1,271,0,5,1,0,0,2] rc2_subset_1
s: 271
i: [1,273,0,5,1,0,0,255] rc3_subset_1
s: 273
i: [1,217,0,6,2,0,0,255] rc4_subset_1
s: 217
i: [1,223,0,6,2,0,0,2] d2_struct_0
s: 223
i: [1,225,0,6,3,0,0,255] cc7_struct_0
s: 225
i: [1,226,0,6,3,0,0,255] cc5_struct_0
s: 226
i: [1,228,0,6,3,0,0,255] cc2_struct_0
s: 228
i: [1,229,0,6,3,0,0,255] cc17_struct_0
s: 229
d: [2,357,1,4,2,1,40,226,333]
a: [3,363,1,2,1,357]
s: 363
i: [1,230,0,6,3,0,0,255] cc6_struct_0
s: 230
d: [2,359,2,4,2,1,40,226,351]
a: [3,372,2,2,1,359]
s: 372
i: [1,231,0,6,3,0,0,255] cc1_struct_0
s: 231
i: [1,234,0,6,3,0,0,255] cc4_struct_0
s: 234
i: [1,243,0,6,2,0,0,2] t3_xboole_1
s: 243
i: [1,262,0,6,2,0,0,255] t1_subset
s: 262
i: [1,276,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 276
i: [1,207,0,7,2,0,0,255] rc6_subset_1
s: 207
i: [1,211,0,7,2,0,0,2] rc1_subset_1
s: 211
i: [1,213,0,7,2,0,0,255] rc5_subset_1
s: 213
i: [1,216,0,7,2,0,0,255] rc4_subset_1
s: 216
i: [1,232,0,7,3,0,0,255] cc8_struct_0
s: 232
i: [1,235,0,7,3,0,0,255] cc11_struct_0
s: 235
i: [1,236,0,7,3,0,0,255] cc11_struct_0
s: 236
i: [1,237,0,7,3,0,0,255] cc9_struct_0
s: 237
i: [1,244,0,7,3,0,0,255] fc9_struct_0
s: 244
i: [1,245,0,7,3,0,0,255] fc6_struct_0
s: 245
i: [1,247,0,7,3,0,0,255] rc22_struct_0
s: 247
i: [1,251,0,7,3,0,0,1] fc2_struct_0
s: 251
i: [1,253,0,7,3,0,0,1] rc4_struct_0
s: 253
i: [1,255,0,7,3,0,0,255] rc21_struct_0
s: 255
i: [1,256,0,7,3,0,0,255] rc21_struct_0
s: 256
i: [1,259,0,7,3,0,0,255] fc8_struct_0
s: 259
i: [1,260,0,7,3,0,0,255] fc7_struct_0
s: 260
i: [1,261,0,7,3,0,0,1] fc1_struct_0
s: 261
i: [1,264,0,7,2,0,0,2] t3_subset
s: 264
d: [2,403,1,4,1,1,40,264,293]
a: [3,414,1,4,1,403]
s: 414
i: [1,265,0,7,2,0,0,2] t3_subset
s: 265
i: [1,266,0,7,3,0,0,1] t8_boole
s: 266
d: [2,409,1,4,1,0,40,264,271]
s: 409
i: [1,220,0,8,2,0,0,1] rc1_connsp_1
s: 220
d: [2,410,1,4,1,0,40,264,273]
s: 410
i: [1,224,0,8,2,0,0,2] dt_k1_struct_0
s: 224
i: [1,239,0,8,3,0,0,255] cc5_subset_1
s: 239
d: [2,435,1,5,2,1,40,239,293]
a: [3,447,1,3,1,435]
s: 447
i: [1,241,0,8,3,0,0,2] cc1_subset_1
s: 241
d: [2,453,2,4,2,1,40,447,260]
d: [2,454,2,2,1,2,21,453,346]
a: [3,458,2,2,1,454]
s: 458
i: [1,263,0,8,3,0,0,255] t2_subset
s: 263
i: [1,238,0,9,4,0,0,255] cc10_struct_0
s: 238
d: [2,462,1,5,2,1,40,241,293]
a: [3,474,1,3,1,462]
s: 474
i: [1,242,0,9,3,0,0,255] cc4_subset_1
s: 242
d: [2,482,1,9,2,1,40,263,293]
d: [2,493,1,5,1,1,21,482,206]
a: [3,497,1,5,1,493]
s: 497
i: [1,250,0,9,4,0,0,255] rc7_pre_topc
s: 250
i: [1,269,0,9,3,0,0,255] t5_subset
s: 269
d: [2,522,2,5,1,1,40,497,276]
a: [3,528,2,5,1,522]
s: 528
i: [1,246,0,10,3,0,0,255] rc22_struct_0
s: 246
d: [2,417,2,4,1,0,40,409,243]
a: [3,421,2,4,1,417]
s: 421
i: [1,252,0,10,3,0,0,1] rc4_struct_0
s: 252
i: [1,254,0,10,3,0,0,255] rc21_struct_0
s: 254
d: [2,424,2,4,1,0,40,410,243]
a: [3,428,2,4,1,424]
s: 428
i: [1,258,0,10,3,0,0,255] rc6_pre_topc
s: 258
d: [2,574,3,3,1,1,44,274,428]
a: [3,579,3,3,1,574]
s: 579
i: [1,268,0,10,3,0,0,255] t4_subset
s: 268
i: [1,209,0,11,4,0,0,255] cc3_subset_1
s: 209
a: [3,477,1,2,1,462]
s: 477
i: [1,210,0,11,4,0,0,255] cc2_subset_1
s: 210
d: [2,591,1,7,2,1,40,268,293]
s: 591
i: [1,219,0,12,4,0,0,1] cc4_connsp_1
s: 219
i: [1,249,0,12,4,0,0,255] rc7_pre_topc
s: 249
d: [2,632,2,10,3,1,40,591,263]
d: [2,633,2,7,2,2,21,632,474]
s: 633
i: [1,248,0,13,4,0,0,1] t3_connsp_3
s: 248
d: [2,667,3,7,2,2,40,633,276]
s: 667
i: [1,277,0,14,3,0,0,1] dt_k1_connsp_3
s: 277
i: [1,257,0,25,7,0,0,1] t1_connsp_3
s: 257
d: [2,672,4,8,2,2,40,667,633]
d: [2,673,4,4,1,2,18,672]
a: [3,677,4,4,1,673]
s: 677
d: [2,688,1,12,3,1,44,277,283]
d: [2,693,1,7,2,2,21,688,293]
d: [2,694,1,5,1,3,21,693,298]
a: [3,698,1,5,1,694]
s: 698
r: 667
r: 633
r: 474
r: 477
s: 477
s: 474
d: [2,703,1,20,6,1,40,257,293]
d: [2,722,1,17,5,2,21,703,288]
d: [2,723,1,15,4,3,21,722,298]
d: [2,724,1,13,3,4,21,723,303]
s: 724
d: [2,728,2,20,6,1,40,698,257]
d: [2,739,2,18,5,2,21,728,298]
d: [2,740,2,16,4,3,21,739,303]
a: [3,743,2,3,1,740]
s: 743
d: [2,408,1,5,1,0,40,264,270]
s: 408
d: [2,789,2,8,2,5,40,724,698]
a: [3,809,2,3,1,789]
s: 809
d: [2,442,1,5,2,0,40,239,271]
s: 442
d: [2,443,1,5,2,0,40,239,273]
s: 443
d: [2,791,2,8,2,4,40,724,293]
d: [2,825,2,5,1,4,21,791,275]
d: [2,826,2,3,1,5,22,825,283]
a: [3,830,2,3,1,826]
s: 830
r: 477
r: 809
a: [3,478,1,2,1,462]
s: 478
a: [3,808,2,3,1,789]
s: 808
d: [2,606,1,9,3,1,40,209,293]
d: [2,623,1,6,2,2,21,606,474]
a: [3,627,1,4,1,623]
s: 627
d: [2,470,1,5,2,0,40,241,273]
s: 470
d: [2,975,3,3,1,1,40,830,243]
a: [3,980,3,3,1,975]
s: 980
r: 283
r: 288
r: 293
r: 414
r: 478
r: 497
r: 528
r: 591
r: 627
r: 677
r: 724
r: 808
r: 830
d: [2,986,0,5,1,2,23,283,980]
a: [3,1002,0,5,1,986]
s: 1002
d: [2,730,2,9,3,1,40,698,209]
d: [2,752,2,7,2,2,21,730,339]
d: [2,753,2,4,1,3,21,752,474]
a: [3,757,2,4,1,753]
s: 757
d: [2,993,1,7,2,2,23,591,980]
s: 993
d: [2,1018,1,7,2,1,44,1002,240]
s: 1018
d: [2,1023,2,14,5,1,44,1018,248]
d: [2,1026,2,12,4,2,21,1023,298]
d: [2,1027,2,10,3,3,21,1026,303]
d: [2,1028,2,8,2,4,21,1027,308]
a: [3,1032,2,3,1,1028]
s: 1032
d: [2,1042,3,0,0,2,52,346,1032,222]
e: 1042
d: [2,1043,3,2,1,1,40,1032,222]
d: [2,1046,3,0,0,2,21,1043,346]
e: 1046
d: [2,1044,3,4,2,1,44,1032,223]
d: [2,1048,3,2,1,2,21,1044,339]
d: [2,1049,3,0,0,3,21,1048,346]
e: 1049
r: 474
r: 757
r: 1032
a: [3,473,1,3,1,462]
s: 473
d: [2,1078,2,0,0,3,52,308,346,473,251]
e: 1078
d: [2,1079,2,4,2,1,40,473,251]
d: [2,1081,2,2,1,2,21,1079,346]
d: [2,1082,2,0,0,3,21,1081,308]
e: 1082
d: [2,1024,2,14,5,1,44,248,1018]
d: [2,1038,2,12,4,2,21,1024,298]
d: [2,1039,2,10,3,3,21,1038,303]
d: [2,1040,2,8,2,4,21,1039,308]
f: 1082,1046,1040,986,975,826,341,197,198,199,200,201,202
% Refutation found. Thanks to Tanya!
% SZS status Theorem for connsp_3__t4_connsp_3
% SZS output start Proof for connsp_3__t4_connsp_3
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ~(k1_connsp_3(X0,X1) = k1_xboole_0 & v2_connsp_1(X1,X0)))) [input t4_connsp_3]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ~(k1_connsp_3(X0,X1) = k1_xboole_0 & v2_connsp_1(X1,X0)))) [negated conjecture 1]
29. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
39. ! [X0] : ((l1_struct_0(X0) & ~v2_struct_0(X0)) => ~v1_xboole_0(u1_struct_0(X0))) [input fc2_struct_0]
40. ! [X0] : (l1_struct_0(X0) => v1_xboole_0(k1_struct_0(X0))) [input fc3_struct_0]
57. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
59. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (? [X2] : (r1_tarski(X1,X2) & v2_connsp_1(X2,X0) & m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) => r1_tarski(X1,k1_connsp_3(X0,X1))))) [input t1_connsp_3]
62. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => u1_struct_0(X0) = k1_connsp_3(X0,k1_struct_0(X0))) [input t3_connsp_3]
64. ! [X0] : (r1_tarski(X0,k1_xboole_0) => k1_xboole_0 = X0) [input t3_xboole_1]
67. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
77. ! [X0] : r1_tarski(X0,X0) [rectify 57]
82. ? [X0] : (? [X1] : ((k1_connsp_3(X0,X1) = k1_xboole_0 & v2_connsp_1(X1,X0)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & (l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
83. ? [X0] : (? [X1] : (k1_connsp_3(X0,X1) = k1_xboole_0 & v2_connsp_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [flattening 82]
92. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 29]
96. ! [X0] : (v1_xboole_0(k1_struct_0(X0)) | ~l1_struct_0(X0)) [ennf transformation 40]
122. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 67]
125. ! [X0] : (k1_xboole_0 = X0 | ~r1_tarski(X0,k1_xboole_0)) [ennf transformation 64]
132. ! [X0] : (u1_struct_0(X0) = k1_connsp_3(X0,k1_struct_0(X0)) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0))) [ennf transformation 62]
133. ! [X0] : (u1_struct_0(X0) = k1_connsp_3(X0,k1_struct_0(X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [flattening 132]
136. ! [X0] : (~v1_xboole_0(u1_struct_0(X0)) | (~l1_struct_0(X0) | v2_struct_0(X0))) [ennf transformation 39]
137. ! [X0] : (~v1_xboole_0(u1_struct_0(X0)) | ~l1_struct_0(X0) | v2_struct_0(X0)) [flattening 136]
142. ! [X0] : (! [X1] : ((r1_tarski(X1,k1_connsp_3(X0,X1)) | ! [X2] : (~r1_tarski(X1,X2) | ~v2_connsp_1(X2,X0) | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 59]
143. ! [X0] : (! [X1] : (r1_tarski(X1,k1_connsp_3(X0,X1)) | ! [X2] : (~r1_tarski(X1,X2) | ~v2_connsp_1(X2,X0) | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 142]
163. ? [X0] : (? [X1] : (k1_connsp_3(X0,X1) = k1_xboole_0 & v2_connsp_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (? [X1] : (k1_xboole_0 = k1_connsp_3(sK0,X1) & v2_connsp_1(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0)) [choice axiom]
164. ? [X1] : (k1_xboole_0 = k1_connsp_3(sK0,X1) & v2_connsp_1(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => (k1_xboole_0 = k1_connsp_3(sK0,sK1) & v2_connsp_1(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
165. (k1_xboole_0 = k1_connsp_3(sK0,sK1) & v2_connsp_1(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0) [skolemisation 83,164,163]
197. ~v2_struct_0(sK0) [cnf transformation 165] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
198. v2_pre_topc(sK0) [cnf transformation 165] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
199. l1_pre_topc(sK0) [cnf transformation 165] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
200. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 165] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
201. v2_connsp_1(sK1,sK0) [cnf transformation 165] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
202. k1_xboole_0 = k1_connsp_3(sK0,sK1) [cnf transformation 165] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
218. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 92] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
222. v1_xboole_0(k1_struct_0(X0)) | ~l1_struct_0(X0) [cnf transformation 96] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
240. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 122] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
243. ~r1_tarski(X0,k1_xboole_0) | k1_xboole_0 = X0 [cnf transformation 125] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
248. u1_struct_0(X0) = k1_connsp_3(X0,k1_struct_0(X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 133] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
251. ~v1_xboole_0(u1_struct_0(X0)) | ~l1_struct_0(X0) | v2_struct_0(X0) [cnf transformation 137] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
257. ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) | ~r1_tarski(X1,X2) | ~v2_connsp_1(X2,X0) | r1_tarski(X1,k1_connsp_3(X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 143] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
275. r1_tarski(X0,X0) [cnf transformation 77] {a:0,w:3,wCS:1200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
281. 1 <=> k1_xboole_0 = k1_connsp_3(sK0,sK1) [avatar definition]
283. k1_xboole_0 = k1_connsp_3(sK0,sK1) <- (1) [avatar component clause 281] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
284. 1 [avatar split clause 202,281]
286. 2 <=> v2_connsp_1(sK1,sK0) [avatar definition]
288. v2_connsp_1(sK1,sK0) <- (2) [avatar component clause 286] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
289. 2 [avatar split clause 201,286]
291. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
293. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (3) [avatar component clause 291] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
294. 3 [avatar split clause 200,291]
296. 4 <=> l1_pre_topc(sK0) [avatar definition]
298. l1_pre_topc(sK0) <- (4) [avatar component clause 296] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
299. 4 [avatar split clause 199,296]
301. 5 <=> v2_pre_topc(sK0) [avatar definition]
303. v2_pre_topc(sK0) <- (5) [avatar component clause 301] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
304. 5 [avatar split clause 198,301]
306. 6 <=> v2_struct_0(sK0) [avatar definition]
308. ~v2_struct_0(sK0) <- (~6) [avatar component clause 306] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
309. ~6 [avatar split clause 197,306]
341. l1_struct_0(sK0) <- (4) [resolution 218,298] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
344. 13 <=> l1_struct_0(sK0) [avatar definition]
346. l1_struct_0(sK0) <- (13) [avatar component clause 344] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
347. 13 | ~4 [avatar split clause 341,296,344]
472. 27 <=> v1_xboole_0(u1_struct_0(sK0)) [avatar definition]
473. v1_xboole_0(u1_struct_0(sK0)) <- (27) [avatar component clause 472] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
703. ~r1_tarski(X8,sK1) | ~v2_connsp_1(sK1,sK0) | r1_tarski(X8,k1_connsp_3(sK0,X8)) | ~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (3) [resolution 257,293] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
722. ~r1_tarski(X8,sK1) | r1_tarski(X8,k1_connsp_3(sK0,X8)) | ~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (2, 3) [subsumption resolution 703,288] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
723. ~r1_tarski(X8,sK1) | r1_tarski(X8,k1_connsp_3(sK0,X8)) | ~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(sK0))) | ~v2_pre_topc(sK0) <- (2, 3, 4) [subsumption resolution 722,298] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
724. ~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(sK0))) | r1_tarski(X8,k1_connsp_3(sK0,X8)) | ~r1_tarski(X8,sK1) <- (2, 3, 4, 5) [subsumption resolution 723,303] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
791. r1_tarski(sK1,k1_connsp_3(sK0,sK1)) | ~r1_tarski(sK1,sK1) <- (2, 3, 4, 5) [resolution 724,293] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
825. r1_tarski(sK1,k1_connsp_3(sK0,sK1)) <- (2, 3, 4, 5) [subsumption resolution 791,275] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
826. r1_tarski(sK1,k1_xboole_0) <- (1, 2, 3, 4, 5) [forward demodulation 825,283] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
828. 48 <=> r1_tarski(sK1,k1_xboole_0) [avatar definition]
830. r1_tarski(sK1,k1_xboole_0) <- (48) [avatar component clause 828] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
831. 48 | ~1 | ~2 | ~3 | ~4 | ~5 [avatar split clause 826,301,296,291,286,281,828]
975. k1_xboole_0 = sK1 <- (48) [resolution 830,243] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
978. 76 <=> k1_xboole_0 = sK1 [avatar definition]
980. k1_xboole_0 = sK1 <- (76) [avatar component clause 978] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
981. 76 | ~48 [avatar split clause 975,828,978]
986. k1_xboole_0 = k1_connsp_3(sK0,k1_xboole_0) <- (1, 76) [backward demodulation 283,980] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1000. 77 <=> k1_xboole_0 = k1_connsp_3(sK0,k1_xboole_0) [avatar definition]
1002. k1_xboole_0 = k1_connsp_3(sK0,k1_xboole_0) <- (77) [avatar component clause 1000] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1003. 77 | ~1 | ~76 [avatar split clause 986,978,281,1000]
1018. k1_connsp_3(sK0,X0) = X0 | ~v1_xboole_0(X0) <- (77) [superposition 1002,240] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1024. u1_struct_0(sK0) = k1_struct_0(sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) | ~v1_xboole_0(k1_struct_0(sK0)) <- (77) [superposition 248,1018] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1030. 80 <=> v1_xboole_0(k1_struct_0(sK0)) [avatar definition]
1032. ~v1_xboole_0(k1_struct_0(sK0)) <- (~80) [avatar component clause 1030] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1034. 81 <=> u1_struct_0(sK0) = k1_struct_0(sK0) [avatar definition]
1038. u1_struct_0(sK0) = k1_struct_0(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) | ~v1_xboole_0(k1_struct_0(sK0)) <- (4, 77) [subsumption resolution 1024,298] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1039. u1_struct_0(sK0) = k1_struct_0(sK0) | v2_struct_0(sK0) | ~v1_xboole_0(k1_struct_0(sK0)) <- (4, 5, 77) [subsumption resolution 1038,303] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1040. u1_struct_0(sK0) = k1_struct_0(sK0) | ~v1_xboole_0(k1_struct_0(sK0)) <- (4, 5, ~6, 77) [subsumption resolution 1039,308] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1041. ~80 | 81 | ~4 | ~5 | 6 | ~77 [avatar split clause 1040,1000,306,301,296,1034,1030]
1043. ~l1_struct_0(sK0) <- (~80) [resolution 1032,222] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1046. $false <- (13, ~80) [subsumption resolution 1043,346] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1047. ~13 | 80 [avatar contradiction clause 1046]
1051. u1_struct_0(sK0) != k1_struct_0(sK0) | ~v1_xboole_0(k1_struct_0(sK0)) | v1_xboole_0(u1_struct_0(sK0)) [theory tautology sat conflict] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
1079. ~l1_struct_0(sK0) | v2_struct_0(sK0) <- (27) [resolution 473,251] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1081. v2_struct_0(sK0) <- (13, 27) [subsumption resolution 1079,346] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1082. $false <- (~6, 13, 27) [subsumption resolution 1081,308] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1083. 6 | ~13 | ~27 [avatar contradiction clause 1082]
1084. $false [avatar sat refutation 284,289,294,299,304,309,347,831,981,1003,1041,1047,1051,1083] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:689,thDist:-689}
% SZS output end Proof for connsp_3__t4_connsp_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 68
% >>> Preprocessing
% Introduced skolems: 17
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 83
% Generated clauses: 551
% Active clauses: 118
% Passive clauses: 291
% Final active clauses: 95
% Final passive clauses: 163
% Main loop iterations started: 123
% 
% >>> Simplifying Inferences
% Duplicate literals: 23
% Fw subsumption resolutions: 76
% Fw demodulations: 6
% Bw demodulations: 14
% 
% >>> Deletion Inferences
% Simple tautologies: 25
% Equational tautologies: 1
% Forward subsumptions: 54
% 
% >>> Generating Inferences
% Binary resolution: 231
% Unit resulting resolution: 2
% Forward superposition: 22
% Backward superposition: 11
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 36
% Split components: 76
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 87
% SAT solver unit clauses: 14
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 5628
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% resolution: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.007 s ( own 0.003 s ) 

