i: [1,68,0,2,1,1,0,0] t31_pboole
a: [3,113,0,2,1,68]
s: 113
i: [1,66,0,2,1,1,0,0] t31_pboole
a: [3,123,0,2,1,66]
s: 123
i: [1,64,0,2,1,1,0,0] t31_pboole
a: [3,133,0,2,1,64]
s: 133
i: [1,62,0,2,1,1,0,0] t31_pboole
a: [3,143,0,2,1,62]
s: 143
i: [1,69,0,3,1,1,0,0] t31_pboole
a: [3,108,0,3,1,69]
s: 108
i: [1,67,0,3,1,1,0,0] t31_pboole
a: [3,118,0,3,1,67]
s: 118
i: [1,65,0,3,1,1,0,0] t31_pboole
a: [3,128,0,3,1,65]
s: 128
i: [1,63,0,3,1,1,0,0] t31_pboole
a: [3,138,0,3,1,63]
s: 138
i: [1,97,0,2,1,0,0,1] rc1_funct_1
a: [3,148,0,2,1,97]
s: 148
i: [1,96,0,2,1,0,0,1] rc1_funct_1
a: [3,153,0,2,1,96]
s: 153
i: [1,70,0,10,1,1,0,0] t31_pboole
a: [3,103,0,10,1,70]
s: 103
i: [1,75,0,3,1,0,0,1] rc1_pboole
s: 75
i: [1,77,0,3,1,0,0,1] rc1_pboole
s: 77
i: [1,76,0,4,1,0,0,1] rc1_pboole
s: 76
i: [1,78,0,4,1,0,0,1] rc1_pboole
s: 78
i: [1,98,0,24,9,0,0,1] redefinition_r6_pboole
d: [2,99,0,14,5,0,18,98]
s: 99
i: [1,79,0,24,9,0,0,1] reflexivity_r6_pboole
s: 79
i: [1,80,0,24,9,0,0,2] reflexivity_r2_pboole
s: 80
i: [1,85,0,25,9,0,0,1] dt_k3_pboole
s: 85
i: [1,87,0,25,9,0,0,1] dt_k3_pboole
s: 87
i: [1,89,0,25,9,0,0,1] dt_k2_pboole
s: 89
i: [1,91,0,25,9,0,0,1] dt_k2_pboole
s: 91
i: [1,81,0,26,9,0,0,1] idempotence_k3_pboole
s: 81
i: [1,82,0,26,9,0,0,1] idempotence_k2_pboole
s: 82
i: [1,86,0,26,9,0,0,1] dt_k3_pboole
s: 86
i: [1,88,0,26,9,0,0,1] dt_k3_pboole
s: 88
i: [1,90,0,26,9,0,0,1] dt_k2_pboole
s: 90
i: [1,92,0,26,9,0,0,1] dt_k2_pboole
s: 92
i: [1,71,0,27,9,0,0,2] t14_pboole
s: 71
i: [1,72,0,27,9,0,0,1] t15_pboole
s: 72
i: [1,94,0,27,10,0,0,1] redefinition_r6_pboole
s: 94
i: [1,93,0,28,10,0,0,1] symmetry_r6_pboole
s: 93
i: [1,83,0,29,9,0,0,1] commutativity_k3_pboole
s: 83
d: [2,184,1,47,17,0,44,72,83]
d: [2,193,1,27,9,0,18,184]
s: 193
i: [1,84,0,29,9,0,0,1] commutativity_k2_pboole
s: 84
i: [1,73,0,31,11,0,0,2] l22_pboole
s: 73
d: [2,179,1,47,17,0,44,72,83]
d: [2,198,1,27,9,0,18,179]
s: 198
i: [1,74,0,45,15,0,0,2] t16_pboole
s: 74
d: [2,210,1,47,17,0,44,71,84]
d: [2,219,1,27,9,0,18,210]
s: 219
d: [2,205,1,47,17,0,44,71,84]
d: [2,224,1,27,9,0,18,205]
s: 224
d: [2,226,1,65,18,0,40,73,71]
d: [2,231,1,55,14,0,18,226]
d: [2,233,1,49,13,0,21,231,92]
d: [2,234,1,44,12,0,21,233,91]
d: [2,235,1,38,11,0,21,234,90]
d: [2,236,1,33,10,0,21,235,89]
s: 236
d: [2,269,2,64,23,0,40,236,74]
d: [2,275,2,34,11,0,18,269]
d: [2,276,2,30,10,0,21,275,80]
s: 276
d: [2,227,1,65,18,0,40,73,72]
d: [2,230,1,55,14,0,18,227]
d: [2,237,1,49,13,0,21,230,88]
d: [2,238,1,44,12,0,21,237,87]
d: [2,239,1,38,11,0,21,238,86]
d: [2,240,1,33,10,0,21,239,85]
s: 240
d: [2,279,3,43,10,1,44,103,276]
d: [2,301,3,40,9,2,21,279,128]
d: [2,302,3,38,8,3,21,301,133]
d: [2,303,3,35,7,4,21,302,138]
d: [2,304,3,33,6,5,21,303,143]
a: [3,308,3,7,1,304]
s: 308
d: [2,340,4,0,0,9,52,143,123,133,113,138,128,118,108,308,72]
e: 340
d: [2,341,4,20,8,1,40,308,72]
d: [2,343,4,17,7,2,21,341,108]
d: [2,344,4,15,6,3,21,343,113]
d: [2,345,4,12,5,4,21,344,118]
d: [2,346,4,10,4,5,21,345,123]
d: [2,347,4,7,3,6,21,346,128]
d: [2,348,4,5,2,7,21,347,133]
d: [2,349,4,2,1,8,21,348,138]
d: [2,350,4,0,0,9,21,349,143]
e: 350
r: 308
d: [2,245,1,83,24,0,40,74,73]
d: [2,252,1,73,20,0,18,245]
d: [2,253,1,67,19,0,21,252,92]
d: [2,254,1,62,18,0,21,253,91]
d: [2,255,1,56,17,0,21,254,90]
d: [2,256,1,51,16,0,21,255,89]
s: 256
d: [2,353,2,64,23,0,40,256,219]
d: [2,367,2,34,11,0,18,353]
d: [2,369,2,30,10,0,21,367,80]
s: 369
a: [3,312,3,5,1,304]
s: 312
d: [2,438,4,0,0,9,52,143,123,133,113,138,128,118,108,312,85]
e: 438
d: [2,439,4,20,8,1,40,312,85]
d: [2,441,4,17,7,2,21,439,108]
d: [2,442,4,15,6,3,21,441,113]
d: [2,443,4,12,5,4,21,442,118]
d: [2,444,4,10,4,5,21,443,123]
d: [2,445,4,7,3,6,21,444,128]
d: [2,446,4,5,2,7,21,445,133]
d: [2,447,4,2,1,8,21,446,138]
d: [2,448,4,0,0,9,21,447,143]
e: 448
r: 312
d: [2,228,2,65,18,0,40,73,193]
d: [2,229,2,55,14,0,18,228]
d: [2,241,2,49,13,0,21,229,88]
d: [2,242,2,44,12,0,21,241,87]
d: [2,243,2,38,11,0,21,242,86]
d: [2,244,2,33,10,0,21,243,85]
s: 244
a: [3,311,3,5,1,304]
s: 311
d: [2,257,2,65,18,0,40,219,73]
d: [2,264,2,55,14,0,18,257]
d: [2,265,2,49,13,0,21,264,92]
d: [2,266,2,44,12,0,21,265,91]
d: [2,267,2,38,11,0,21,266,90]
d: [2,268,2,33,10,0,21,267,89]
s: 268
d: [2,272,2,53,18,0,44,236,84]
d: [2,273,2,33,10,0,18,272]
s: 273
a: [3,316,3,6,1,304]
s: 316
d: [2,470,4,0,0,9,52,143,123,133,113,138,128,118,108,316,86]
e: 470
d: [2,471,4,20,8,1,40,316,86]
d: [2,473,4,17,7,2,21,471,108]
d: [2,474,4,15,6,3,21,473,113]
d: [2,475,4,12,5,4,21,474,118]
d: [2,476,4,10,4,5,21,475,123]
d: [2,477,4,7,3,6,21,476,128]
d: [2,478,4,5,2,7,21,477,133]
d: [2,479,4,2,1,8,21,478,138]
d: [2,480,4,0,0,9,21,479,143]
e: 480
r: 316
d: [2,271,2,53,18,0,44,236,84]
d: [2,274,2,33,10,0,18,271]
s: 274
a: [3,320,3,5,1,304]
s: 320
d: [2,482,4,0,0,9,52,143,123,133,113,138,128,118,108,320,87]
e: 482
d: [2,483,4,20,8,1,40,320,87]
d: [2,485,4,17,7,2,21,483,108]
d: [2,486,4,15,6,3,21,485,113]
d: [2,487,4,12,5,4,21,486,118]
d: [2,488,4,10,4,5,21,487,123]
d: [2,489,4,7,3,6,21,488,128]
d: [2,490,4,5,2,7,21,489,133]
d: [2,491,4,2,1,8,21,490,138]
d: [2,492,4,0,0,9,21,491,143]
e: 492
r: 320
d: [2,333,2,53,18,0,44,240,83]
d: [2,334,2,33,10,0,18,333]
s: 334
d: [2,332,2,53,18,0,44,240,83]
d: [2,335,2,33,10,0,18,332]
s: 335
a: [3,319,3,5,1,304]
s: 319
d: [2,355,2,118,23,0,40,256,72]
d: [2,370,2,109,22,0,21,355,88]
d: [2,371,2,101,21,0,21,370,87]
d: [2,372,2,92,20,0,21,371,86]
d: [2,373,2,84,19,0,21,372,85]
d: [2,374,2,78,18,0,21,373,92]
d: [2,375,2,73,17,0,21,374,91]
d: [2,376,2,67,16,0,21,375,90]
d: [2,377,2,62,15,0,21,376,89]
s: 377
a: [3,315,3,6,1,304]
s: 315
d: [2,356,2,118,23,0,40,256,193]
d: [2,378,2,109,22,0,21,356,88]
d: [2,379,2,101,21,0,21,378,87]
d: [2,380,2,92,20,0,21,379,86]
d: [2,381,2,84,19,0,21,380,85]
d: [2,382,2,78,18,0,21,381,92]
d: [2,383,2,73,17,0,21,382,91]
d: [2,384,2,67,16,0,21,383,90]
d: [2,385,2,62,15,0,21,384,89]
s: 385
d: [2,357,2,121,29,0,40,256,74]
d: [2,386,2,115,28,0,21,357,92]
d: [2,387,2,110,27,0,21,386,91]
d: [2,388,2,104,26,0,21,387,90]
d: [2,389,2,99,25,0,21,388,89]
d: [2,390,2,93,24,0,21,389,92]
d: [2,391,2,88,23,0,21,390,91]
d: [2,392,2,82,22,0,21,391,90]
d: [2,393,2,77,21,0,21,392,89]
s: 393
a: [3,324,3,6,1,304]
s: 324
d: [2,735,4,0,0,9,52,143,123,133,113,138,128,118,108,324,88]
e: 735
d: [2,736,4,20,8,1,40,324,88]
d: [2,738,4,17,7,2,21,736,108]
d: [2,739,4,15,6,3,21,738,113]
d: [2,740,4,12,5,4,21,739,118]
d: [2,741,4,10,4,5,21,740,123]
d: [2,742,4,7,3,6,21,741,128]
d: [2,743,4,5,2,7,21,742,133]
d: [2,744,4,2,1,8,21,743,138]
d: [2,745,4,0,0,9,21,744,143]
e: 745
r: 324
d: [2,330,2,118,23,0,40,240,74]
d: [2,336,2,112,22,0,21,330,92]
d: [2,337,2,107,21,0,21,336,91]
d: [2,338,2,101,20,0,21,337,90]
d: [2,339,2,96,19,0,21,338,89]
s: 339
a: [3,328,3,4,1,304]
s: 328
d: [2,747,4,0,0,5,52,143,133,138,128,328,99]
e: 747
d: [2,748,4,10,4,1,40,328,99]
d: [2,750,4,7,3,2,21,748,128]
d: [2,751,4,5,2,3,21,750,133]
d: [2,752,4,2,1,4,21,751,138]
d: [2,753,4,0,0,5,21,752,143]
e: 753
f: 753,745,492,480,448,350,304,62,63,64,65,66,67,68,69,70
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pboole__t31_pboole
% SZS output start Proof for pboole__t31_pboole
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r6_pboole(X0,k2_pboole(X0,X1,k3_pboole(X0,X1,X2)),X1))) [input t31_pboole]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r6_pboole(X0,k2_pboole(X0,X1,k3_pboole(X0,X1,X2)),X1))) [negated conjecture 1]
5. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2)))) [input dt_k2_pboole]
6. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2)))) [input dt_k3_pboole]
9. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => X1 = X2))) [input l22_pboole]
12. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (r6_pboole(X0,X1,X2) <=> X1 = X2)) [input redefinition_r6_pboole]
13. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => r2_pboole(X0,X1,X1)) [input reflexivity_r2_pboole]
16. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,X1,k2_pboole(X0,X1,X2)))) [input t14_pboole]
17. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k3_pboole(X0,X1,X2),X1))) [input t15_pboole]
18. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X3,X2) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,k2_pboole(X0,X1,X3),X2))))) [input t16_pboole]
19. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r6_pboole(X0,k2_pboole(X0,X1,k3_pboole(X0,X1,X2)),X1))) [flattening 2]
20. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,X1,k2_pboole(X0,X1,X2)))) [flattening 16]
21. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k3_pboole(X0,X1,X2),X1))) [flattening 17]
22. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => X1 = X2))) [flattening 9]
23. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X3,X2) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,k2_pboole(X0,X1,X3),X2))))) [flattening 18]
24. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k2_pboole(X0,X1,k3_pboole(X0,X1,X2)),X1) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 19]
25. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k2_pboole(X0,X1,k3_pboole(X0,X1,X2)),X1) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 24]
26. ! [X0,X1] : (! [X2] : (r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 20]
27. ! [X0,X1] : (! [X2] : (r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 26]
28. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 21]
29. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 28]
30. ! [X0,X1] : (! [X2] : ((X1 = X2 | (~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 22]
31. ! [X0,X1] : (! [X2] : (X1 = X2 | ~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 30]
32. ! [X0,X1] : (! [X2] : (! [X3] : ((r2_pboole(X0,k2_pboole(X0,X1,X3),X2) | (~r2_pboole(X0,X3,X2) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 23]
33. ! [X0,X1] : (! [X2] : (! [X3] : (r2_pboole(X0,k2_pboole(X0,X1,X3),X2) | ~r2_pboole(X0,X3,X2) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 32]
36. ! [X0,X1,X2] : (r2_pboole(X0,X1,X1) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 13]
37. ! [X0,X1,X2] : (r2_pboole(X0,X1,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 36]
46. ! [X0,X1,X2] : ((v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 6]
47. ! [X0,X1,X2] : ((v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 46]
48. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 5]
49. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 48]
52. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 12]
53. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 52]
54. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,k2_pboole(X0,X1,k3_pboole(X0,X1,X2)),X1) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (~r6_pboole(sK0,k2_pboole(sK0,sK1,k3_pboole(sK0,sK1,X2)),sK1) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
55. ? [X2] : (~r6_pboole(sK0,k2_pboole(sK0,sK1,k3_pboole(sK0,sK1,X2)),sK1) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (~r6_pboole(sK0,k2_pboole(sK0,sK1,k3_pboole(sK0,sK1,sK2)),sK1) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
56. (~r6_pboole(sK0,k2_pboole(sK0,sK1,k3_pboole(sK0,sK1,sK2)),sK1) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 25,55,54]
59. ! [X0,X1,X2] : (((r6_pboole(X0,X1,X2) | X1 != X2) & (X1 = X2 | ~r6_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 53]
62. v1_relat_1(sK1) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
63. v4_relat_1(sK1,sK0) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
64. v1_funct_1(sK1) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. v1_partfun1(sK1,sK0) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
66. v1_relat_1(sK2) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
67. v4_relat_1(sK2,sK0) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
68. v1_funct_1(sK2) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
69. v1_partfun1(sK2,sK0) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
70. ~r6_pboole(sK0,k2_pboole(sK0,sK1,k3_pboole(sK0,sK1,sK2)),sK1) [cnf transformation 56] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
71. r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 27] {a:0,w:27,wCS:10800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
72. r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 29] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. ~r2_pboole(X0,X2,X1) | X1 = X2 | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 31] {a:0,w:31,wCS:12400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
74. r2_pboole(X0,k2_pboole(X0,X1,X3),X2) | ~r2_pboole(X0,X3,X2) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 33] {a:0,w:45,wCS:18000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
80. r2_pboole(X0,X1,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 37] {a:0,w:24,wCS:9600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
85. v1_relat_1(k3_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 47] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
86. v4_relat_1(k3_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 47] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. v1_funct_1(k3_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 47] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
88. v1_partfun1(k3_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 47] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
89. v1_relat_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 49] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
90. v4_relat_1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 49] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
91. v1_funct_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 49] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
92. v1_partfun1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 49] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
95. r6_pboole(X0,X1,X2) | X1 != X2 | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 59] {a:0,w:27,wCS:10800,sine:1,thAx:0,allAx:1,thDist:-1}
98. r6_pboole(X0,X2,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) [equality resolution 95] {a:0,w:24,wCS:9600,sine:1,thAx:0,allAx:1,thDist:-1}
99. r6_pboole(X0,X2,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) [duplicate literal removal 98] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. 1 <=> r6_pboole(sK0,k2_pboole(sK0,sK1,k3_pboole(sK0,sK1,sK2)),sK1) [avatar definition]
103. ~r6_pboole(sK0,k2_pboole(sK0,sK1,k3_pboole(sK0,sK1,sK2)),sK1) <- (~1) [avatar component clause 101] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. ~1 [avatar split clause 70,101]
106. 2 <=> v1_partfun1(sK2,sK0) [avatar definition]
108. v1_partfun1(sK2,sK0) <- (2) [avatar component clause 106] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
109. 2 [avatar split clause 69,106]
111. 3 <=> v1_funct_1(sK2) [avatar definition]
113. v1_funct_1(sK2) <- (3) [avatar component clause 111] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 3 [avatar split clause 68,111]
116. 4 <=> v4_relat_1(sK2,sK0) [avatar definition]
118. v4_relat_1(sK2,sK0) <- (4) [avatar component clause 116] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. 4 [avatar split clause 67,116]
121. 5 <=> v1_relat_1(sK2) [avatar definition]
123. v1_relat_1(sK2) <- (5) [avatar component clause 121] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. 5 [avatar split clause 66,121]
126. 6 <=> v1_partfun1(sK1,sK0) [avatar definition]
128. v1_partfun1(sK1,sK0) <- (6) [avatar component clause 126] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. 6 [avatar split clause 65,126]
131. 7 <=> v1_funct_1(sK1) [avatar definition]
133. v1_funct_1(sK1) <- (7) [avatar component clause 131] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. 7 [avatar split clause 64,131]
136. 8 <=> v4_relat_1(sK1,sK0) [avatar definition]
138. v4_relat_1(sK1,sK0) <- (8) [avatar component clause 136] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
139. 8 [avatar split clause 63,136]
141. 9 <=> v1_relat_1(sK1) [avatar definition]
143. v1_relat_1(sK1) <- (9) [avatar component clause 141] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. 9 [avatar split clause 62,141]
226. k2_pboole(X3,X4,X5) = X4 | ~r2_pboole(X3,k2_pboole(X3,X4,X5),X4) | ~v1_partfun1(X4,X3) | ~v1_funct_1(X4) | ~v4_relat_1(X4,X3) | ~v1_relat_1(X4) | ~v1_partfun1(k2_pboole(X3,X4,X5),X3) | ~v1_funct_1(k2_pboole(X3,X4,X5)) | ~v4_relat_1(k2_pboole(X3,X4,X5),X3) | ~v1_relat_1(k2_pboole(X3,X4,X5)) | ~v1_partfun1(X5,X3) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X3) | ~v1_relat_1(X5) | ~v1_partfun1(X4,X3) | ~v1_funct_1(X4) | ~v4_relat_1(X4,X3) | ~v1_relat_1(X4) [resolution 73,71] {a:1,w:65,wCS:26000,sine:2,thAx:0,allAx:2,thDist:-2}
231. k2_pboole(X3,X4,X5) = X4 | ~r2_pboole(X3,k2_pboole(X3,X4,X5),X4) | ~v1_partfun1(X4,X3) | ~v1_funct_1(X4) | ~v4_relat_1(X4,X3) | ~v1_relat_1(X4) | ~v1_partfun1(k2_pboole(X3,X4,X5),X3) | ~v1_funct_1(k2_pboole(X3,X4,X5)) | ~v4_relat_1(k2_pboole(X3,X4,X5),X3) | ~v1_relat_1(k2_pboole(X3,X4,X5)) | ~v1_partfun1(X5,X3) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X3) | ~v1_relat_1(X5) [duplicate literal removal 226] {a:1,w:55,wCS:22000,sine:2,thAx:0,allAx:2,thDist:-2}
233. k2_pboole(X3,X4,X5) = X4 | ~r2_pboole(X3,k2_pboole(X3,X4,X5),X4) | ~v1_partfun1(X4,X3) | ~v1_funct_1(X4) | ~v4_relat_1(X4,X3) | ~v1_relat_1(X4) | ~v1_funct_1(k2_pboole(X3,X4,X5)) | ~v4_relat_1(k2_pboole(X3,X4,X5),X3) | ~v1_relat_1(k2_pboole(X3,X4,X5)) | ~v1_partfun1(X5,X3) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X3) | ~v1_relat_1(X5) [subsumption resolution 231,92] {a:1,w:49,wCS:19600,sine:1,thAx:0,allAx:2,thDist:-2}
234. k2_pboole(X3,X4,X5) = X4 | ~r2_pboole(X3,k2_pboole(X3,X4,X5),X4) | ~v1_partfun1(X4,X3) | ~v1_funct_1(X4) | ~v4_relat_1(X4,X3) | ~v1_relat_1(X4) | ~v4_relat_1(k2_pboole(X3,X4,X5),X3) | ~v1_relat_1(k2_pboole(X3,X4,X5)) | ~v1_partfun1(X5,X3) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X3) | ~v1_relat_1(X5) [subsumption resolution 233,91] {a:1,w:44,wCS:17600,sine:1,thAx:0,allAx:2,thDist:-2}
235. k2_pboole(X3,X4,X5) = X4 | ~r2_pboole(X3,k2_pboole(X3,X4,X5),X4) | ~v1_partfun1(X4,X3) | ~v1_funct_1(X4) | ~v4_relat_1(X4,X3) | ~v1_relat_1(X4) | ~v1_relat_1(k2_pboole(X3,X4,X5)) | ~v1_partfun1(X5,X3) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X3) | ~v1_relat_1(X5) [subsumption resolution 234,90] {a:1,w:38,wCS:15200,sine:1,thAx:0,allAx:2,thDist:-2}
236. ~r2_pboole(X3,k2_pboole(X3,X4,X5),X4) | k2_pboole(X3,X4,X5) = X4 | ~v1_partfun1(X4,X3) | ~v1_funct_1(X4) | ~v4_relat_1(X4,X3) | ~v1_relat_1(X4) | ~v1_partfun1(X5,X3) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X3) | ~v1_relat_1(X5) [subsumption resolution 235,89] {a:1,w:33,wCS:13200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
269. k2_pboole(X0,X1,X2) = X1 | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [resolution 236,74] {a:2,w:64,wCS:25600,sine:1,thAx:0,allAx:3,thDist:-3}
275. k2_pboole(X0,X1,X2) = X1 | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X1) [duplicate literal removal 269] {a:2,w:34,wCS:13600,sine:1,thAx:0,allAx:3,thDist:-3}
276. k2_pboole(X0,X1,X2) = X1 | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~r2_pboole(X0,X2,X1) [subsumption resolution 275,80] {a:2,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
279. ~r6_pboole(sK0,sK1,sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) <- (~1) [superposition 103,276] {a:3,w:43,wCS:4300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
301. ~r6_pboole(sK0,sK1,sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) <- (~1, 6) [subsumption resolution 279,128] {a:3,w:40,wCS:4000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
302. ~r6_pboole(sK0,sK1,sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) <- (~1, 6, 7) [subsumption resolution 301,133] {a:3,w:38,wCS:3800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
303. ~r6_pboole(sK0,sK1,sK1) | ~v1_relat_1(sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) <- (~1, 6, 7, 8) [subsumption resolution 302,138] {a:3,w:35,wCS:3500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
304. ~r6_pboole(sK0,sK1,sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) <- (~1, 6, 7, 8, 9) [subsumption resolution 303,143] {a:3,w:33,wCS:3300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
306. 12 <=> r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) [avatar definition]
308. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) <- (~12) [avatar component clause 306] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
310. 13 <=> v1_relat_1(k3_pboole(sK0,sK1,sK2)) [avatar definition]
312. ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) <- (~13) [avatar component clause 310] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
314. 14 <=> v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) [avatar definition]
316. ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) <- (~14) [avatar component clause 314] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
318. 15 <=> v1_funct_1(k3_pboole(sK0,sK1,sK2)) [avatar definition]
320. ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) <- (~15) [avatar component clause 318] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
322. 16 <=> v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) [avatar definition]
324. ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) <- (~16) [avatar component clause 322] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
326. 17 <=> r6_pboole(sK0,sK1,sK1) [avatar definition]
328. ~r6_pboole(sK0,sK1,sK1) <- (~17) [avatar component clause 326] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
329. ~12 | ~13 | ~14 | ~15 | ~16 | ~17 | 1 | ~6 | ~7 | ~8 | ~9 [avatar split clause 304,141,136,131,126,101,326,322,318,314,310,306]
341. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~12) [resolution 308,72] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
343. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~12) [subsumption resolution 341,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
344. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~12) [subsumption resolution 343,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
345. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~12) [subsumption resolution 344,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
346. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~12) [subsumption resolution 345,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
347. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~12) [subsumption resolution 346,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
348. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~12) [subsumption resolution 347,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
349. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~12) [subsumption resolution 348,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
350. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~12) [subsumption resolution 349,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
351. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 12 [avatar contradiction clause 350]
439. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~13) [resolution 312,85] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
441. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~13) [subsumption resolution 439,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
442. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~13) [subsumption resolution 441,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
443. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~13) [subsumption resolution 442,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
444. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~13) [subsumption resolution 443,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
445. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~13) [subsumption resolution 444,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
446. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~13) [subsumption resolution 445,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
447. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~13) [subsumption resolution 446,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
448. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~13) [subsumption resolution 447,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
449. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 13 [avatar contradiction clause 448]
471. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~14) [resolution 316,86] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
473. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~14) [subsumption resolution 471,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
474. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~14) [subsumption resolution 473,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
475. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~14) [subsumption resolution 474,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
476. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~14) [subsumption resolution 475,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
477. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~14) [subsumption resolution 476,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
478. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~14) [subsumption resolution 477,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
479. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~14) [subsumption resolution 478,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
480. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~14) [subsumption resolution 479,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
481. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 14 [avatar contradiction clause 480]
483. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~15) [resolution 320,87] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
485. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~15) [subsumption resolution 483,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
486. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~15) [subsumption resolution 485,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
487. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~15) [subsumption resolution 486,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
488. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~15) [subsumption resolution 487,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
489. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~15) [subsumption resolution 488,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
490. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~15) [subsumption resolution 489,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
491. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~15) [subsumption resolution 490,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
492. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~15) [subsumption resolution 491,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
493. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 15 [avatar contradiction clause 492]
736. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~16) [resolution 324,88] {a:4,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
738. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, ~16) [subsumption resolution 736,108] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
739. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, ~16) [subsumption resolution 738,113] {a:4,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
740. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, ~16) [subsumption resolution 739,118] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
741. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, ~16) [subsumption resolution 740,123] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
742. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, ~16) [subsumption resolution 741,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
743. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, ~16) [subsumption resolution 742,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
744. ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, ~16) [subsumption resolution 743,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
745. $false <- (2, 3, 4, 5, 6, 7, 8, 9, ~16) [subsumption resolution 744,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
746. ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 16 [avatar contradiction clause 745]
748. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~17) [resolution 328,99] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
750. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, ~17) [subsumption resolution 748,128] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
751. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, ~17) [subsumption resolution 750,133] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
752. ~v1_relat_1(sK1) <- (6, 7, 8, ~17) [subsumption resolution 751,138] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
753. $false <- (6, 7, 8, 9, ~17) [subsumption resolution 752,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
754. ~6 | ~7 | ~8 | ~9 | 17 [avatar contradiction clause 753]
755. $false [avatar sat refutation 104,109,114,119,124,129,134,139,144,329,351,449,481,493,746,754] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:325,thDist:-325}
% SZS output end Proof for pboole__t31_pboole
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 602
% Active clauses: 56
% Passive clauses: 111
% Final active clauses: 51
% Final passive clauses: 47
% Main loop iterations started: 64
% 
% >>> Simplifying Inferences
% Duplicate literals: 1284
% Fw subsumption resolutions: 176
% 
% >>> Deletion Inferences
% Simple tautologies: 33
% Equational tautologies: 3
% Forward subsumptions: 88
% 
% >>> Generating Inferences
% Binary resolution: 24
% Unit resulting resolution: 6
% Forward superposition: 58
% Backward superposition: 102
% Self superposition: 4
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 16
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 29
% SAT solver unit clauses: 11
% 
% Memory used [KB]: 5628
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% forward subsumption index maintenance: 0.001 s
% forward subsumption: 0.002 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.001 s
% superposition: 0.001 s
% other: 0.008 s ( own 0.001 s ) 

