i: [1,70,0,2,1,1,0,0] l1_polyred
a: [3,103,0,2,1,70]
s: 103
i: [1,68,0,2,1,1,0,0] l1_polyred
a: [3,113,0,2,1,68]
s: 113
i: [1,67,0,2,1,1,0,0] l1_polyred
a: [3,118,0,2,1,67]
s: 118
i: [1,66,0,2,1,1,0,0] l1_polyred
a: [3,123,0,2,1,66]
s: 123
i: [1,71,0,3,1,1,0,0] l1_polyred
a: [3,98,0,3,1,71]
s: 98
i: [1,65,0,3,1,1,0,0] l1_polyred
a: [3,128,0,3,1,65]
s: 128
i: [1,64,0,4,1,1,0,0] l1_polyred
a: [3,133,0,4,1,64]
s: 133
i: [1,69,0,6,1,1,0,0] l1_polyred
a: [3,108,0,6,1,69]
s: 108
i: [1,72,0,3,1,0,0,255] fc1_finsub_1
s: 72
i: [1,84,0,3,1,0,0,2] reflexivity_r1_tarski
s: 84
i: [1,73,0,4,2,0,0,255] cc1_finsub_1
s: 73
d: [2,135,1,3,1,0,40,73,72]
s: 135
i: [1,74,0,4,2,0,0,255] cc1_finsub_1
s: 74
d: [2,136,1,3,1,0,40,74,72]
s: 136
i: [1,83,0,4,1,0,0,1] existence_m1_subset_1
s: 83
i: [1,85,0,4,1,0,0,1] fc6_relat_1
s: 85
i: [1,92,0,4,1,0,0,255] rc3_relat_1
s: 92
i: [1,93,0,5,1,0,0,255] rc3_relat_1
s: 93
i: [1,94,0,5,1,0,0,255] rc3_relat_1
s: 94
i: [1,78,0,6,3,0,0,255] cc2_finsub_1
s: 78
i: [1,81,0,7,2,0,0,2] t3_subset
s: 81
d: [2,138,1,3,1,1,40,81,133]
a: [3,144,1,3,1,138]
s: 144
i: [1,82,0,7,2,0,0,2] t3_subset
s: 82
d: [2,139,1,5,1,1,40,81,108]
a: [3,149,1,5,1,139]
s: 149
i: [1,75,0,8,3,0,0,1] t37_orders_1
s: 75
i: [1,76,0,8,3,0,0,2] cc2_relat_1
s: 76
d: [2,152,1,4,2,1,40,76,133]
a: [3,159,1,2,1,152]
s: 159
i: [1,88,0,8,2,0,0,1] cc1_relset_1
s: 88
d: [2,153,1,6,2,1,40,76,108]
d: [2,165,1,2,1,1,21,153,85]
a: [3,169,1,2,1,165]
s: 169
i: [1,89,0,9,2,0,0,255] cc2_relset_1
s: 89
i: [1,90,0,9,2,0,0,255] cc2_relset_1
s: 90
d: [2,175,1,3,1,1,40,89,108]
a: [3,181,1,3,1,175]
s: 181
i: [1,77,0,11,4,0,0,1] t38_orders_1
s: 77
d: [2,183,1,3,1,1,40,90,108]
a: [3,189,1,3,1,183]
s: 189
i: [1,86,0,12,4,0,0,255] cc5_relat_1
s: 86
i: [1,87,0,12,4,0,0,255] cc6_relat_1
s: 87
d: [2,195,1,12,3,1,40,86,108]
d: [2,198,1,8,2,1,21,195,85]
s: 198
i: [1,79,0,13,3,0,0,255] cc4_funct_2
s: 79
d: [2,200,1,12,3,1,40,87,108]
d: [2,203,1,8,2,1,21,200,85]
s: 203
i: [1,91,0,13,3,0,0,255] cc1_funct_2
s: 91
i: [1,80,0,19,6,0,0,255] t15_orders_1
s: 80
d: [2,207,1,7,2,1,40,91,108]
d: [2,210,1,4,1,2,21,207,128]
a: [3,214,1,4,1,210]
s: 214
d: [2,216,1,13,5,1,40,80,108]
d: [2,219,1,10,4,2,21,216,128]
d: [2,220,1,8,3,3,21,219,113]
d: [2,221,1,6,2,4,21,220,118]
d: [2,222,1,4,1,5,21,221,123]
a: [3,226,1,4,1,222]
s: 226
d: [2,228,2,7,3,1,44,75,226]
d: [2,229,2,5,2,2,21,228,103]
d: [2,230,2,3,1,3,21,229,169]
a: [3,234,2,3,1,230]
s: 234
d: [2,236,3,0,0,4,52,169,144,98,234,77]
e: 236
d: [2,171,1,2,1,1,40,88,108]
f: 236,230,222,171,138,64,65,66,67,68,69,70,71
% Refutation found. Thanks to Tanya!
% SZS status Theorem for polyred__l1_polyred
% SZS output start Proof for polyred__l1_polyred
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) => (v3_orders_1(X2) => r3_orders_1(X2,X1)))) [input l1_polyred]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) => (v3_orders_1(X2) => r3_orders_1(X2,X1)))) [negated conjecture 1]
5. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [input cc1_relset_1]
21. ! [X0] : ! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v8_relat_2(X1) & v4_relat_2(X1) & v1_relat_2(X1)) => k1_relat_1(X1) = X0) [input t15_orders_1]
22. ! [X0] : (v1_relat_1(X0) => (v3_orders_1(X0) => r3_orders_1(X0,k1_relat_1(X0)))) [input t37_orders_1]
23. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r3_orders_1(X0,X1)) => r3_orders_1(X0,X2))) [input t38_orders_1]
24. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
25. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) => (v3_orders_1(X2) => r3_orders_1(X2,X1)))) [flattening 2]
26. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : ((r1_tarski(X2,X1) & r3_orders_1(X0,X1)) => r3_orders_1(X0,X2))) [flattening 23]
28. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v8_relat_2(X1) & v4_relat_2(X1) & v1_relat_2(X1)) => k1_relat_1(X1) = X0) [flattening 21]
29. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 24]
31. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [flattening 5]
34. ? [X0,X1] : (? [X2] : ((~r3_orders_1(X2,X1) & v3_orders_1(X2)) & (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 25]
35. ? [X0,X1] : (? [X2] : (~r3_orders_1(X2,X1) & v3_orders_1(X2) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) & m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 34]
37. ! [X0] : ((r3_orders_1(X0,k1_relat_1(X0)) | ~v3_orders_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 22]
38. ! [X0] : (r3_orders_1(X0,k1_relat_1(X0)) | ~v3_orders_1(X0) | ~v1_relat_1(X0)) [flattening 37]
40. ! [X0] : (! [X1,X2] : (r3_orders_1(X0,X2) | (~r1_tarski(X2,X1) | ~r3_orders_1(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 26]
41. ! [X0] : (! [X1,X2] : (r3_orders_1(X0,X2) | ~r1_tarski(X2,X1) | ~r3_orders_1(X0,X1)) | ~v1_relat_1(X0)) [flattening 40]
46. ! [X0,X1] : (k1_relat_1(X1) = X0 | (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~v1_partfun1(X1,X0) | ~v8_relat_2(X1) | ~v4_relat_2(X1) | ~v1_relat_2(X1))) [ennf transformation 28]
47. ! [X0,X1] : (k1_relat_1(X1) = X0 | ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~v1_partfun1(X1,X0) | ~v8_relat_2(X1) | ~v4_relat_2(X1) | ~v1_relat_2(X1)) [flattening 46]
52. ! [X0,X1,X2] : (v1_relat_1(X2) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 31]
56. ? [X0,X1] : (? [X2] : (~r3_orders_1(X2,X1) & v3_orders_1(X2) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,X0)) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (? [X2] : (~r3_orders_1(X2,sK1) & v3_orders_1(X2) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,sK0)) & m1_subset_1(sK1,k1_zfmisc_1(sK0))) [choice axiom]
57. ? [X2] : (~r3_orders_1(X2,sK1) & v3_orders_1(X2) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v8_relat_2(X2) & v4_relat_2(X2) & v1_relat_2(X2) & v1_partfun1(X2,sK0)) => (~r3_orders_1(sK2,sK1) & v3_orders_1(sK2) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v8_relat_2(sK2) & v4_relat_2(sK2) & v1_relat_2(sK2) & v1_partfun1(sK2,sK0)) [choice axiom]
58. (~r3_orders_1(sK2,sK1) & v3_orders_1(sK2) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v8_relat_2(sK2) & v4_relat_2(sK2) & v1_relat_2(sK2) & v1_partfun1(sK2,sK0)) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) [skolemisation 35,57,56]
59. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 29]
64. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 58] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. v1_partfun1(sK2,sK0) [cnf transformation 58] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
66. v1_relat_2(sK2) [cnf transformation 58] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
67. v4_relat_2(sK2) [cnf transformation 58] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
68. v8_relat_2(sK2) [cnf transformation 58] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
69. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) [cnf transformation 58] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
70. v3_orders_1(sK2) [cnf transformation 58] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
71. ~r3_orders_1(sK2,sK1) [cnf transformation 58] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
75. r3_orders_1(X0,k1_relat_1(X0)) | ~v3_orders_1(X0) | ~v1_relat_1(X0) [cnf transformation 38] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
77. ~r1_tarski(X2,X1) | r3_orders_1(X0,X2) | ~r3_orders_1(X0,X1) | ~v1_relat_1(X0) [cnf transformation 41] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
80. ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | k1_relat_1(X1) = X0 | ~v1_partfun1(X1,X0) | ~v8_relat_2(X1) | ~v4_relat_2(X1) | ~v1_relat_2(X1) [cnf transformation 47] {a:0,w:19,wCS:7600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
81. ~m1_subset_1(X0,k1_zfmisc_1(X1)) | r1_tarski(X0,X1) [cnf transformation 59] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
88. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v1_relat_1(X2) [cnf transformation 52] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
96. 1 <=> r3_orders_1(sK2,sK1) [avatar definition]
98. ~r3_orders_1(sK2,sK1) <- (~1) [avatar component clause 96] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
99. ~1 [avatar split clause 71,96]
101. 2 <=> v3_orders_1(sK2) [avatar definition]
103. v3_orders_1(sK2) <- (2) [avatar component clause 101] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. 2 [avatar split clause 70,101]
106. 3 <=> m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) [avatar definition]
108. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) <- (3) [avatar component clause 106] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
109. 3 [avatar split clause 69,106]
111. 4 <=> v8_relat_2(sK2) [avatar definition]
113. v8_relat_2(sK2) <- (4) [avatar component clause 111] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 4 [avatar split clause 68,111]
116. 5 <=> v4_relat_2(sK2) [avatar definition]
118. v4_relat_2(sK2) <- (5) [avatar component clause 116] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. 5 [avatar split clause 67,116]
121. 6 <=> v1_relat_2(sK2) [avatar definition]
123. v1_relat_2(sK2) <- (6) [avatar component clause 121] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. 6 [avatar split clause 66,121]
126. 7 <=> v1_partfun1(sK2,sK0) [avatar definition]
128. v1_partfun1(sK2,sK0) <- (7) [avatar component clause 126] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. 7 [avatar split clause 65,126]
131. 8 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
133. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (8) [avatar component clause 131] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. 8 [avatar split clause 64,131]
138. r1_tarski(sK1,sK0) <- (8) [resolution 81,133] {a:1,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
142. 9 <=> r1_tarski(sK1,sK0) [avatar definition]
144. r1_tarski(sK1,sK0) <- (9) [avatar component clause 142] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
145. 9 | ~8 [avatar split clause 138,131,142]
167. 13 <=> v1_relat_1(sK2) [avatar definition]
169. v1_relat_1(sK2) <- (13) [avatar component clause 167] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
171. v1_relat_1(sK2) <- (3) [resolution 88,108] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
174. 13 | ~3 [avatar split clause 171,106,167]
216. sK0 = k1_relat_1(sK2) | ~v1_partfun1(sK2,sK0) | ~v8_relat_2(sK2) | ~v4_relat_2(sK2) | ~v1_relat_2(sK2) <- (3) [resolution 80,108] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
219. sK0 = k1_relat_1(sK2) | ~v8_relat_2(sK2) | ~v4_relat_2(sK2) | ~v1_relat_2(sK2) <- (3, 7) [subsumption resolution 216,128] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
220. sK0 = k1_relat_1(sK2) | ~v4_relat_2(sK2) | ~v1_relat_2(sK2) <- (3, 4, 7) [subsumption resolution 219,113] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
221. sK0 = k1_relat_1(sK2) | ~v1_relat_2(sK2) <- (3, 4, 5, 7) [subsumption resolution 220,118] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
222. sK0 = k1_relat_1(sK2) <- (3, 4, 5, 6, 7) [subsumption resolution 221,123] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
224. 17 <=> sK0 = k1_relat_1(sK2) [avatar definition]
226. sK0 = k1_relat_1(sK2) <- (17) [avatar component clause 224] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
227. 17 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 222,126,121,116,111,106,224]
228. r3_orders_1(sK2,sK0) | ~v3_orders_1(sK2) | ~v1_relat_1(sK2) <- (17) [superposition 75,226] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
229. r3_orders_1(sK2,sK0) | ~v1_relat_1(sK2) <- (2, 17) [subsumption resolution 228,103] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
230. r3_orders_1(sK2,sK0) <- (2, 13, 17) [subsumption resolution 229,169] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
232. 18 <=> r3_orders_1(sK2,sK0) [avatar definition]
234. r3_orders_1(sK2,sK0) <- (18) [avatar component clause 232] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
235. 18 | ~2 | ~13 | ~17 [avatar split clause 230,224,167,101,232]
236. $false <- (~1, 9, 13, 18) [unit resulting resolution 169,144,98,234,77] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:17,thDist:-17}
237. 1 | ~9 | ~13 | ~18 [avatar contradiction clause 236]
238. $false [avatar sat refutation 99,104,109,114,119,124,129,134,145,174,227,235,237] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:118,thDist:-118}
% SZS output end Proof for polyred__l1_polyred
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 31
% Generated clauses: 100
% Active clauses: 44
% Passive clauses: 65
% Final active clauses: 44
% Final passive clauses: 21
% Main loop iterations started: 44
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 40
% Unit resulting resolution: 1
% Backward superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 8
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

