i: [1,48,0,2,1,1,0,0] t46_algspec1
a: [3,80,0,2,1,48]
s: 80
i: [1,47,0,2,1,1,0,0] t46_algspec1
a: [3,85,0,2,1,47]
s: 85
i: [1,51,0,3,1,1,0,0] t46_algspec1
a: [3,65,0,3,1,51]
s: 65
i: [1,50,0,3,1,1,0,0] t46_algspec1
a: [3,70,0,3,1,50]
s: 70
i: [1,49,0,3,1,1,0,0] t46_algspec1
a: [3,75,0,3,1,49]
s: 75
i: [1,61,0,2,1,0,0,1] existence_l1_msualg_1
a: [3,90,0,2,1,61]
s: 90
i: [1,59,0,8,3,0,0,2] existence_m1_instalg1
s: 59
i: [1,60,0,8,3,0,0,1] existence_m2_algspec1
s: 60
i: [1,52,0,9,4,0,0,2] cc6_instalg1
s: 52
i: [1,53,0,9,4,0,0,2] dt_m1_instalg1
s: 53
i: [1,55,0,9,4,0,0,1] dt_m2_algspec1
s: 55
d: [2,96,1,6,3,1,40,55,70]
a: [3,102,1,2,1,96]
s: 102
i: [1,56,0,9,4,0,0,1] dt_m2_algspec1
s: 56
a: [3,106,1,2,1,96]
s: 106
i: [1,54,0,13,5,0,0,2] t16_instalg1
s: 54
i: [1,57,0,14,6,0,0,1] d5_algspec1
s: 57
a: [3,111,1,2,1,96]
s: 111
d: [2,132,0,12,5,0,21,57,55]
d: [2,133,0,10,4,0,21,132,56]
s: 133
d: [2,116,1,6,3,1,40,56,70]
d: [2,120,1,4,2,2,21,116,102]
a: [3,124,1,2,1,120]
s: 124
i: [1,58,0,14,6,0,0,1] d5_algspec1
s: 58
d: [2,152,0,12,5,0,21,58,52]
d: [2,153,0,10,4,0,21,152,53]
s: 153
d: [2,134,1,7,3,1,40,133,70]
d: [2,138,1,5,2,2,21,134,106]
d: [2,139,1,3,1,3,21,138,102]
a: [3,143,1,3,1,139]
s: 143
d: [2,135,1,7,3,1,40,133,75]
d: [2,145,1,5,2,2,21,135,80]
d: [2,146,1,3,1,3,21,145,85]
a: [3,150,1,3,1,146]
s: 150
d: [2,157,2,10,4,1,40,143,54]
s: 157
d: [2,92,1,11,5,0,40,52,59]
d: [2,93,1,7,3,0,18,92]
s: 93
d: [2,94,1,11,5,0,40,53,59]
d: [2,95,1,7,3,0,18,94]
s: 95
d: [2,161,2,10,4,1,40,150,54]
s: 161
d: [2,98,1,11,5,0,40,55,60]
d: [2,99,1,7,3,0,18,98]
s: 99
d: [2,164,3,7,3,2,40,161,143]
d: [2,165,3,5,2,3,21,164,124]
d: [2,166,3,3,1,4,21,165,111]
a: [3,170,3,3,1,166]
s: 170
d: [2,172,4,0,0,4,52,111,124,65,170,153]
e: 172
d: [2,173,4,7,3,1,40,170,153]
d: [2,178,4,4,2,2,21,173,65]
d: [2,179,4,2,1,3,21,178,124]
d: [2,180,4,0,0,4,21,179,111]
e: 180
d: [2,97,1,6,3,1,40,55,75]
d: [2,113,1,4,2,2,21,97,80]
d: [2,114,1,2,1,3,21,113,85]
d: [2,117,1,6,3,1,40,56,75]
d: [2,126,1,4,2,2,21,117,80]
d: [2,127,1,2,1,3,21,126,85]
f: 180,166,146,139,127,120,114,96,47,48,49,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for algspec1__t46_algspec1
% SZS output start Proof for algspec1__t46_algspec1
1. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : (m2_algspec1(X1,X0) => ! [X2] : (m2_algspec1(X2,X1) => m2_algspec1(X2,X0)))) [input t46_algspec1]
2. ~! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : (m2_algspec1(X1,X0) => ! [X2] : (m2_algspec1(X2,X1) => m2_algspec1(X2,X0)))) [negated conjecture 1]
3. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : (m1_instalg1(X1,X0) => v1_instalg1(X1))) [input cc6_instalg1]
4. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : ((l1_msualg_1(X1) & v1_instalg1(X1)) => (m2_algspec1(X1,X0) <=> m1_instalg1(X0,X1)))) [input d5_algspec1]
8. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : (m1_instalg1(X1,X0) => l1_msualg_1(X1))) [input dt_m1_instalg1]
9. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : (m2_algspec1(X1,X0) => (l1_msualg_1(X1) & v1_instalg1(X1)))) [input dt_m2_algspec1]
15. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : (m1_instalg1(X1,X0) => ! [X2] : (m1_instalg1(X2,X1) => m1_instalg1(X2,X0)))) [input t16_instalg1]
20. ? [X0] : (? [X1] : (? [X2] : (~m2_algspec1(X2,X0) & m2_algspec1(X2,X1)) & m2_algspec1(X1,X0)) & (l1_msualg_1(X0) & v1_instalg1(X0))) [ennf transformation 2]
21. ? [X0] : (? [X1] : (? [X2] : (~m2_algspec1(X2,X0) & m2_algspec1(X2,X1)) & m2_algspec1(X1,X0)) & l1_msualg_1(X0) & v1_instalg1(X0)) [flattening 20]
22. ! [X0] : (! [X1] : (v1_instalg1(X1) | ~m1_instalg1(X1,X0)) | (~l1_msualg_1(X0) | ~v1_instalg1(X0))) [ennf transformation 3]
23. ! [X0] : (! [X1] : (v1_instalg1(X1) | ~m1_instalg1(X1,X0)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [flattening 22]
24. ! [X0] : (! [X1] : (l1_msualg_1(X1) | ~m1_instalg1(X1,X0)) | (~l1_msualg_1(X0) | ~v1_instalg1(X0))) [ennf transformation 8]
25. ! [X0] : (! [X1] : (l1_msualg_1(X1) | ~m1_instalg1(X1,X0)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [flattening 24]
26. ! [X0] : (! [X1] : (! [X2] : (m1_instalg1(X2,X0) | ~m1_instalg1(X2,X1)) | ~m1_instalg1(X1,X0)) | (~l1_msualg_1(X0) | ~v1_instalg1(X0))) [ennf transformation 15]
27. ! [X0] : (! [X1] : (! [X2] : (m1_instalg1(X2,X0) | ~m1_instalg1(X2,X1)) | ~m1_instalg1(X1,X0)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [flattening 26]
28. ! [X0] : (! [X1] : ((l1_msualg_1(X1) & v1_instalg1(X1)) | ~m2_algspec1(X1,X0)) | (~l1_msualg_1(X0) | ~v1_instalg1(X0))) [ennf transformation 9]
29. ! [X0] : (! [X1] : ((l1_msualg_1(X1) & v1_instalg1(X1)) | ~m2_algspec1(X1,X0)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [flattening 28]
30. ! [X0] : (! [X1] : ((m2_algspec1(X1,X0) <=> m1_instalg1(X0,X1)) | (~l1_msualg_1(X1) | ~v1_instalg1(X1))) | (~l1_msualg_1(X0) | ~v1_instalg1(X0))) [ennf transformation 4]
31. ! [X0] : (! [X1] : ((m2_algspec1(X1,X0) <=> m1_instalg1(X0,X1)) | ~l1_msualg_1(X1) | ~v1_instalg1(X1)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [flattening 30]
36. ? [X0] : (? [X1] : (? [X2] : (~m2_algspec1(X2,X0) & m2_algspec1(X2,X1)) & m2_algspec1(X1,X0)) & l1_msualg_1(X0) & v1_instalg1(X0)) => (? [X1] : (? [X2] : (~m2_algspec1(X2,sK0) & m2_algspec1(X2,X1)) & m2_algspec1(X1,sK0)) & l1_msualg_1(sK0) & v1_instalg1(sK0)) [choice axiom]
37. ? [X1] : (? [X2] : (~m2_algspec1(X2,sK0) & m2_algspec1(X2,X1)) & m2_algspec1(X1,sK0)) => (? [X2] : (~m2_algspec1(X2,sK0) & m2_algspec1(X2,sK1)) & m2_algspec1(sK1,sK0)) [choice axiom]
38. ? [X2] : (~m2_algspec1(X2,sK0) & m2_algspec1(X2,sK1)) => (~m2_algspec1(sK2,sK0) & m2_algspec1(sK2,sK1)) [choice axiom]
39. ((~m2_algspec1(sK2,sK0) & m2_algspec1(sK2,sK1)) & m2_algspec1(sK1,sK0)) & l1_msualg_1(sK0) & v1_instalg1(sK0) [skolemisation 21,38,37,36]
40. ! [X0] : (! [X1] : (((m2_algspec1(X1,X0) | ~m1_instalg1(X0,X1)) & (m1_instalg1(X0,X1) | ~m2_algspec1(X1,X0))) | ~l1_msualg_1(X1) | ~v1_instalg1(X1)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [nnf transformation 31]
47. v1_instalg1(sK0) [cnf transformation 39] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
48. l1_msualg_1(sK0) [cnf transformation 39] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
49. m2_algspec1(sK1,sK0) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
50. m2_algspec1(sK2,sK1) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
51. ~m2_algspec1(sK2,sK0) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
52. ~m1_instalg1(X1,X0) | v1_instalg1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 23] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
53. ~m1_instalg1(X1,X0) | l1_msualg_1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 25] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
54. ~m1_instalg1(X2,X1) | m1_instalg1(X2,X0) | ~m1_instalg1(X1,X0) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 27] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
55. ~m2_algspec1(X1,X0) | v1_instalg1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 29] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. ~m2_algspec1(X1,X0) | l1_msualg_1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 29] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. m1_instalg1(X0,X1) | ~m2_algspec1(X1,X0) | ~l1_msualg_1(X1) | ~v1_instalg1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 40] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:1,thDist:-1}
58. m2_algspec1(X1,X0) | ~m1_instalg1(X0,X1) | ~l1_msualg_1(X1) | ~v1_instalg1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 40] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:1,thDist:-1}
63. 1 <=> m2_algspec1(sK2,sK0) [avatar definition]
65. ~m2_algspec1(sK2,sK0) <- (~1) [avatar component clause 63] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66. ~1 [avatar split clause 51,63]
68. 2 <=> m2_algspec1(sK2,sK1) [avatar definition]
70. m2_algspec1(sK2,sK1) <- (2) [avatar component clause 68] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
71. 2 [avatar split clause 50,68]
73. 3 <=> m2_algspec1(sK1,sK0) [avatar definition]
75. m2_algspec1(sK1,sK0) <- (3) [avatar component clause 73] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
76. 3 [avatar split clause 49,73]
78. 4 <=> l1_msualg_1(sK0) [avatar definition]
80. l1_msualg_1(sK0) <- (4) [avatar component clause 78] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
81. 4 [avatar split clause 48,78]
83. 5 <=> v1_instalg1(sK0) [avatar definition]
85. v1_instalg1(sK0) <- (5) [avatar component clause 83] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
86. 5 [avatar split clause 47,83]
96. v1_instalg1(sK2) | ~l1_msualg_1(sK1) | ~v1_instalg1(sK1) <- (2) [resolution 55,70] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
97. v1_instalg1(sK1) | ~l1_msualg_1(sK0) | ~v1_instalg1(sK0) <- (3) [resolution 55,75] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
101. 7 <=> v1_instalg1(sK1) [avatar definition]
102. v1_instalg1(sK1) <- (7) [avatar component clause 101] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
105. 8 <=> l1_msualg_1(sK1) [avatar definition]
106. l1_msualg_1(sK1) <- (8) [avatar component clause 105] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
109. 9 <=> v1_instalg1(sK2) [avatar definition]
111. v1_instalg1(sK2) <- (9) [avatar component clause 109] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
112. ~7 | ~8 | 9 | ~2 [avatar split clause 96,68,109,105,101]
113. v1_instalg1(sK1) | ~v1_instalg1(sK0) <- (3, 4) [subsumption resolution 97,80] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
114. v1_instalg1(sK1) <- (3, 4, 5) [subsumption resolution 113,85] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
115. 7 | ~3 | ~4 | ~5 [avatar split clause 114,83,78,73,101]
116. l1_msualg_1(sK2) | ~l1_msualg_1(sK1) | ~v1_instalg1(sK1) <- (2) [resolution 56,70] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
117. l1_msualg_1(sK1) | ~l1_msualg_1(sK0) | ~v1_instalg1(sK0) <- (3) [resolution 56,75] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
120. l1_msualg_1(sK2) | ~l1_msualg_1(sK1) <- (2, 7) [subsumption resolution 116,102] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
122. 10 <=> l1_msualg_1(sK2) [avatar definition]
124. l1_msualg_1(sK2) <- (10) [avatar component clause 122] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
125. ~8 | 10 | ~2 | ~7 [avatar split clause 120,101,68,122,105]
126. l1_msualg_1(sK1) | ~v1_instalg1(sK0) <- (3, 4) [subsumption resolution 117,80] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
127. l1_msualg_1(sK1) <- (3, 4, 5) [subsumption resolution 126,85] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
128. 8 | ~3 | ~4 | ~5 [avatar split clause 127,83,78,73,105]
132. m1_instalg1(X0,X1) | ~m2_algspec1(X1,X0) | ~l1_msualg_1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [subsumption resolution 57,55] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
133. ~m2_algspec1(X1,X0) | m1_instalg1(X0,X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [subsumption resolution 132,56] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
134. m1_instalg1(sK1,sK2) | ~l1_msualg_1(sK1) | ~v1_instalg1(sK1) <- (2) [resolution 133,70] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
135. m1_instalg1(sK0,sK1) | ~l1_msualg_1(sK0) | ~v1_instalg1(sK0) <- (3) [resolution 133,75] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
138. m1_instalg1(sK1,sK2) | ~v1_instalg1(sK1) <- (2, 8) [subsumption resolution 134,106] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
139. m1_instalg1(sK1,sK2) <- (2, 7, 8) [subsumption resolution 138,102] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
141. 11 <=> m1_instalg1(sK1,sK2) [avatar definition]
143. m1_instalg1(sK1,sK2) <- (11) [avatar component clause 141] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
144. 11 | ~2 | ~7 | ~8 [avatar split clause 139,105,101,68,141]
145. m1_instalg1(sK0,sK1) | ~v1_instalg1(sK0) <- (3, 4) [subsumption resolution 135,80] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
146. m1_instalg1(sK0,sK1) <- (3, 4, 5) [subsumption resolution 145,85] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
148. 12 <=> m1_instalg1(sK0,sK1) [avatar definition]
150. m1_instalg1(sK0,sK1) <- (12) [avatar component clause 148] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
151. 12 | ~3 | ~4 | ~5 [avatar split clause 146,83,78,73,148]
152. m2_algspec1(X1,X0) | ~m1_instalg1(X0,X1) | ~l1_msualg_1(X1) | ~v1_instalg1(X1) | ~l1_msualg_1(X0) [subsumption resolution 58,52] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
153. ~m1_instalg1(X0,X1) | m2_algspec1(X1,X0) | ~l1_msualg_1(X1) | ~v1_instalg1(X1) [subsumption resolution 152,53] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
161. ~m1_instalg1(sK1,X0) | m1_instalg1(sK0,X0) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) <- (12) [resolution 150,54] {a:2,w:10,wCS:1000,nSel:1,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
164. m1_instalg1(sK0,sK2) | ~l1_msualg_1(sK2) | ~v1_instalg1(sK2) <- (11, 12) [resolution 161,143] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:11,thDist:-11}
165. m1_instalg1(sK0,sK2) | ~v1_instalg1(sK2) <- (10, 11, 12) [subsumption resolution 164,124] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:11,thDist:-11}
166. m1_instalg1(sK0,sK2) <- (9, 10, 11, 12) [subsumption resolution 165,111] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:11,thDist:-11}
168. 13 <=> m1_instalg1(sK0,sK2) [avatar definition]
170. m1_instalg1(sK0,sK2) <- (13) [avatar component clause 168] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
171. 13 | ~9 | ~10 | ~11 | ~12 [avatar split clause 166,148,141,122,109,168]
173. m2_algspec1(sK2,sK0) | ~l1_msualg_1(sK2) | ~v1_instalg1(sK2) <- (13) [resolution 170,153] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
178. ~l1_msualg_1(sK2) | ~v1_instalg1(sK2) <- (~1, 13) [subsumption resolution 173,65] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
179. ~v1_instalg1(sK2) <- (~1, 10, 13) [subsumption resolution 178,124] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
180. $false <- (~1, 9, 10, 13) [subsumption resolution 179,111] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
181. 1 | ~9 | ~10 | ~13 [avatar contradiction clause 180]
182. $false [avatar sat refutation 66,71,76,81,86,112,115,125,128,144,151,171,181] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:136,thDist:-136}
% SZS output end Proof for algspec1__t46_algspec1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 82
% Active clauses: 27
% Passive clauses: 34
% Final active clauses: 27
% Final passive clauses: 4
% Main loop iterations started: 29
% 
% >>> Simplifying Inferences
% Duplicate literals: 12
% Fw subsumption resolutions: 20
% 
% >>> Deletion Inferences
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 26
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 5
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 6
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% non unit clause index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

