i: [1,36,0,2,1,1,0,0] t35_funct_4
a: [3,63,0,2,1,36]
s: 63
i: [1,35,0,2,1,1,0,0] t35_funct_4
a: [3,68,0,2,1,35]
s: 68
i: [1,34,0,2,1,1,0,0] t35_funct_4
a: [3,73,0,2,1,34]
s: 73
i: [1,33,0,2,1,1,0,0] t35_funct_4
a: [3,78,0,2,1,33]
s: 78
i: [1,37,0,5,1,1,0,0] t35_funct_4
a: [3,58,0,5,1,37]
s: 58
i: [1,38,0,7,1,1,0,0] t35_funct_4
a: [3,53,0,7,1,38]
s: 53
i: [1,49,0,2,1,0,0,1] rc1_funct_1
a: [3,94,0,2,1,49]
s: 94
i: [1,48,0,2,1,0,0,1] rc1_funct_1
a: [3,99,0,2,1,48]
s: 99
i: [1,42,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 42
i: [1,43,0,11,5,0,0,2] reflexivity_r1_partfun1
a: [3,85,0,7,3,43]
s: 85
i: [1,44,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,89,0,9,3,44]
s: 89
d: [2,101,1,5,1,1,40,42,58]
a: [3,105,1,5,1,101]
s: 105
i: [1,45,0,12,5,0,0,1] dt_k1_funct_4
s: 45
i: [1,46,0,12,5,0,0,1] dt_k1_funct_4
s: 46
i: [1,47,0,14,6,0,0,2] symmetry_r1_partfun1
s: 47
i: [1,39,0,16,6,0,0,1] t56_partfun1
s: 39
d: [2,114,1,11,5,1,40,39,58]
d: [2,116,1,9,4,2,21,114,63]
d: [2,117,1,7,3,3,21,116,68]
d: [2,118,1,5,2,4,21,117,73]
d: [2,119,1,3,1,5,21,118,78]
a: [3,123,1,3,1,119]
s: 123
i: [1,40,0,18,6,0,0,1] t34_funct_4
s: 40
d: [2,140,2,0,0,6,52,78,73,68,63,123,53,40]
e: 140
d: [2,144,1,18,6,1,44,53,40]
d: [2,149,1,11,5,1,19,144]
d: [2,158,1,8,4,2,21,149,123]
d: [2,159,1,6,3,3,21,158,63]
d: [2,160,1,4,2,4,21,159,68]
d: [2,161,1,2,1,5,21,160,73]
d: [2,162,1,0,0,6,21,161,78]
e: 162
f: 162,119,33,34,35,36,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t35_funct_4
% SZS output start Proof for funct_4__t35_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k1_funct_4(X0,X1) = k1_funct_4(X1,X0)))) [input t35_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k1_funct_4(X0,X1) = k1_funct_4(X1,X0)))) [negated conjecture 1]
10. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) <=> k1_funct_4(X0,X1) = k1_funct_4(X1,X0)))) [input t34_funct_4]
11. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => r1_partfun1(X0,X1)))) [input t56_partfun1]
12. ? [X0] : (? [X1] : ((k1_funct_4(X0,X1) != k1_funct_4(X1,X0) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
13. ? [X0] : (? [X1] : (k1_funct_4(X0,X1) != k1_funct_4(X1,X0) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 12]
14. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 11]
15. ! [X0] : (! [X1] : (r1_partfun1(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 14]
16. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> k1_funct_4(X0,X1) = k1_funct_4(X1,X0)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 10]
17. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> k1_funct_4(X0,X1) = k1_funct_4(X1,X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 16]
27. ? [X0] : (? [X1] : (k1_funct_4(X0,X1) != k1_funct_4(X1,X0) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (k1_funct_4(sK0,X1) != k1_funct_4(X1,sK0) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
28. ? [X1] : (k1_funct_4(sK0,X1) != k1_funct_4(X1,sK0) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) => (k1_funct_4(sK0,sK1) != k1_funct_4(sK1,sK0) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
29. (k1_funct_4(sK0,sK1) != k1_funct_4(sK1,sK0) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 13,28,27]
30. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | k1_funct_4(X0,X1) != k1_funct_4(X1,X0)) & (k1_funct_4(X0,X1) = k1_funct_4(X1,X0) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 17]
33. v1_relat_1(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
34. v1_funct_1(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
35. v1_relat_1(sK1) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
36. v1_funct_1(sK1) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
37. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) [cnf transformation 29] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. k1_funct_4(sK0,sK1) != k1_funct_4(sK1,sK0) [cnf transformation 29] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 15] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
40. k1_funct_4(X0,X1) = k1_funct_4(X1,X0) | ~r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. 1 <=> k1_funct_4(sK0,sK1) = k1_funct_4(sK1,sK0) [avatar definition]
53. k1_funct_4(sK0,sK1) != k1_funct_4(sK1,sK0) <- (~1) [avatar component clause 51] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
54. ~1 [avatar split clause 38,51]
56. 2 <=> r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) [avatar definition]
58. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) <- (2) [avatar component clause 56] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
59. 2 [avatar split clause 37,56]
61. 3 <=> v1_funct_1(sK1) [avatar definition]
63. v1_funct_1(sK1) <- (3) [avatar component clause 61] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
64. 3 [avatar split clause 36,61]
66. 4 <=> v1_relat_1(sK1) [avatar definition]
68. v1_relat_1(sK1) <- (4) [avatar component clause 66] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
69. 4 [avatar split clause 35,66]
71. 5 <=> v1_funct_1(sK0) [avatar definition]
73. v1_funct_1(sK0) <- (5) [avatar component clause 71] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
74. 5 [avatar split clause 34,71]
76. 6 <=> v1_relat_1(sK0) [avatar definition]
78. v1_relat_1(sK0) <- (6) [avatar component clause 76] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
79. 6 [avatar split clause 33,76]
114. r1_partfun1(sK0,sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 39,58] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
116. r1_partfun1(sK0,sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3) [subsumption resolution 114,63] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
117. r1_partfun1(sK0,sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3, 4) [subsumption resolution 116,68] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
118. r1_partfun1(sK0,sK1) | ~v1_relat_1(sK0) <- (2, 3, 4, 5) [subsumption resolution 117,73] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
119. r1_partfun1(sK0,sK1) <- (2, 3, 4, 5, 6) [subsumption resolution 118,78] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
121. 13 <=> r1_partfun1(sK0,sK1) [avatar definition]
123. r1_partfun1(sK0,sK1) <- (13) [avatar component clause 121] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
124. 13 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar split clause 119,76,71,66,61,56,121]
144. k1_funct_4(sK0,sK1) != k1_funct_4(sK0,sK1) | ~r1_partfun1(sK0,sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [superposition 53,40] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
149. ~r1_partfun1(sK0,sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [trivial inequality removal 144] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
158. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 13) [subsumption resolution 149,123] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
159. ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 3, 13) [subsumption resolution 158,63] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
160. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 3, 4, 13) [subsumption resolution 159,68] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
161. ~v1_relat_1(sK0) <- (~1, 3, 4, 5, 13) [subsumption resolution 160,73] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
162. $false <- (~1, 3, 4, 5, 6, 13) [subsumption resolution 161,78] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
163. 1 | ~3 | ~4 | ~5 | ~6 | ~13 [avatar contradiction clause 162]
164. $false [avatar sat refutation 54,59,64,69,74,79,124,163] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:98,thDist:-98}
% SZS output end Proof for funct_4__t35_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 75
% Active clauses: 18
% Passive clauses: 20
% Final active clauses: 18
% Final passive clauses: 2
% Main loop iterations started: 18
% 
% >>> Simplifying Inferences
% Duplicate literals: 28
% Trivial inequalities: 2
% Fw subsumption resolutions: 21
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 6
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 6
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 8
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

