i: [1,233,0,2,1,1,0,0] t63_xreal_1
a: [3,360,0,2,1,233]
s: 360
i: [1,234,0,5,1,1,0,0] t63_xreal_1
a: [3,355,0,5,1,234]
s: 355
i: [1,235,0,2,1,0,0,2] fc1_numbers
a: [3,365,0,2,1,235]
s: 365
i: [1,236,0,2,1,0,0,3] spc1_boole
a: [3,370,0,2,1,236]
s: 370
i: [1,237,0,2,1,0,0,2] fc1_xboole_0
a: [3,375,0,2,1,237]
s: 375
i: [1,246,0,2,1,0,0,3] fc6_ordinal1
a: [3,415,0,2,1,246]
s: 415
i: [1,245,0,2,1,0,0,3] fc6_ordinal1
a: [3,420,0,2,1,245]
s: 420
i: [1,247,0,2,1,0,0,2] spc1_numerals
a: [3,440,0,2,1,247]
s: 440
i: [1,319,0,2,1,0,0,2] rc2_xboole_0
a: [3,452,0,2,1,319]
s: 452
i: [1,320,0,2,1,0,0,4] rc4_ordinal1
a: [3,457,0,2,1,320]
s: 457
i: [1,321,0,2,1,0,0,4] rc1_ordinal1
a: [3,462,0,2,1,321]
s: 462
i: [1,322,0,2,1,0,0,1] rc1_xreal_0
a: [3,467,0,2,1,322]
s: 467
i: [1,323,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,472,0,2,1,323]
s: 472
i: [1,324,0,2,1,0,0,2] rc1_xxreal_0
a: [3,477,0,2,1,324]
s: 477
i: [1,325,0,2,1,0,0,2] rc1_xboole_0
a: [3,482,0,2,1,325]
s: 482
i: [1,329,0,2,1,0,0,255] rc2_ordinal1
a: [3,487,0,2,1,329]
s: 487
i: [1,328,0,2,1,0,0,255] rc2_ordinal1
a: [3,492,0,2,1,328]
s: 492
i: [1,327,0,2,1,0,0,255] rc2_ordinal1
a: [3,497,0,2,1,327]
s: 497
i: [1,326,0,2,1,0,0,255] rc2_ordinal1
a: [3,502,0,2,1,326]
s: 502
i: [1,331,0,2,1,0,0,4] rc5_ordinal1
a: [3,507,0,2,1,331]
s: 507
i: [1,330,0,2,1,0,0,4] rc5_ordinal1
a: [3,512,0,2,1,330]
s: 512
i: [1,333,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,517,0,2,1,333]
s: 517
i: [1,332,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,522,0,2,1,332]
s: 522
i: [1,337,0,2,1,0,0,2] rc3_xreal_0
a: [3,527,0,2,1,337]
s: 527
i: [1,336,0,2,1,0,0,2] rc3_xreal_0
a: [3,532,0,2,1,336]
s: 532
i: [1,335,0,2,1,0,0,2] rc3_xreal_0
a: [3,537,0,2,1,335]
s: 537
i: [1,334,0,2,1,0,0,2] rc3_xreal_0
a: [3,542,0,2,1,334]
s: 542
i: [1,341,0,2,1,0,0,2] rc2_xreal_0
a: [3,547,0,2,1,341]
s: 547
i: [1,340,0,2,1,0,0,2] rc2_xreal_0
a: [3,552,0,2,1,340]
s: 552
i: [1,339,0,2,1,0,0,2] rc2_xreal_0
a: [3,557,0,2,1,339]
s: 557
i: [1,338,0,2,1,0,0,2] rc2_xreal_0
a: [3,562,0,2,1,338]
s: 562
i: [1,343,0,2,1,0,0,2] rc3_xxreal_0
a: [3,567,0,2,1,343]
s: 567
i: [1,342,0,2,1,0,0,2] rc3_xxreal_0
a: [3,572,0,2,1,342]
s: 572
i: [1,345,0,2,1,0,0,2] rc2_xxreal_0
a: [3,577,0,2,1,345]
s: 577
i: [1,344,0,2,1,0,0,2] rc2_xxreal_0
a: [3,582,0,2,1,344]
s: 582
i: [1,349,0,2,1,0,0,2] rc4_xreal_0
a: [3,587,0,2,1,349]
s: 587
i: [1,348,0,2,1,0,0,2] rc4_xreal_0
a: [3,592,0,2,1,348]
s: 592
i: [1,347,0,2,1,0,0,2] rc4_xreal_0
a: [3,597,0,2,1,347]
s: 597
i: [1,346,0,2,1,0,0,2] rc4_xreal_0
a: [3,602,0,2,1,346]
s: 602
i: [1,351,0,2,1,0,0,2] rc4_xxreal_0
a: [3,607,0,2,1,351]
s: 607
i: [1,350,0,2,1,0,0,2] rc4_xxreal_0
a: [3,612,0,2,1,350]
s: 612
i: [1,302,0,3,1,0,0,255] reflexivity_r1_tarski
s: 302
i: [1,238,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,380,0,3,1,238]
s: 380
r: 375
d: [2,614,0,2,1,2,23,375,380]
a: [3,618,0,2,1,614]
s: 618
i: [1,239,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,385,0,3,1,239]
s: 385
i: [1,240,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,390,0,3,1,240]
s: 390
i: [1,241,0,3,1,0,0,2] t1_numerals
a: [3,395,0,3,1,241]
s: 395
i: [1,250,0,3,1,0,0,2] spc1_numerals
a: [3,425,0,3,1,250]
s: 425
i: [1,249,0,3,1,0,0,2] spc1_numerals
a: [3,430,0,3,1,249]
s: 430
d: [2,620,0,3,1,2,22,395,380]
a: [3,624,0,3,1,620]
s: 624
d: [2,626,0,3,1,2,22,430,385]
a: [3,630,0,3,1,626]
s: 630
i: [1,252,0,4,2,0,0,4] cc6_ordinal1
s: 252
i: [1,253,0,4,2,0,0,4] cc2_xreal_0
s: 253
d: [2,632,1,2,1,1,40,252,457]
a: [3,637,1,2,1,632]
s: 637
i: [1,254,0,4,2,0,0,4] cc2_xcmplx_0
s: 254
i: [1,255,0,4,2,0,0,4] cc2_xxreal_0
s: 255
d: [2,633,1,2,1,1,40,252,507]
a: [3,642,1,2,1,633]
s: 642
i: [1,256,0,4,2,0,0,255] cc1_ordinal1
s: 256
d: [2,644,1,2,1,1,40,253,457]
a: [3,649,1,2,1,644]
s: 649
i: [1,257,0,4,2,0,0,255] cc1_ordinal1
s: 257
i: [1,259,0,4,2,0,0,2] cc3_xreal_0
s: 259
d: [2,730,1,2,1,1,40,259,360]
a: [3,739,1,2,1,730]
s: 739
i: [1,260,0,4,2,0,0,2] cc4_xreal_0
s: 260
d: [2,747,1,2,1,1,40,260,360]
a: [3,756,1,2,1,747]
s: 756
i: [1,271,0,4,2,0,0,255] cc9_ordinal1
s: 271
i: [1,272,0,4,2,0,0,4] cc7_ordinal1
s: 272
d: [2,645,1,2,1,1,40,253,507]
a: [3,654,1,2,1,645]
s: 654
i: [1,273,0,4,2,0,0,4] cc3_ordinal1
s: 273
d: [2,656,1,2,1,1,40,254,457]
a: [3,661,1,2,1,656]
s: 661
i: [1,301,0,4,1,0,0,3] existence_m1_subset_1
s: 301
i: [1,242,0,4,1,0,0,2] dt_k5_numbers
a: [3,400,0,4,1,242]
s: 400
d: [2,657,1,2,1,1,40,254,507]
a: [3,666,1,2,1,657]
s: 666
i: [1,243,0,4,1,0,0,1] dt_k6_numbers
a: [3,405,0,4,1,243]
s: 405
d: [2,668,1,2,1,1,40,255,457]
a: [3,673,1,2,1,668]
s: 673
i: [1,248,0,4,1,0,0,2] spc1_numerals
a: [3,435,0,4,1,248]
s: 435
d: [2,840,0,4,1,2,22,400,385]
a: [3,844,0,4,1,840]
s: 844
d: [2,669,1,2,1,1,40,255,507]
a: [3,678,1,2,1,669]
s: 678
d: [2,846,0,4,1,2,22,405,385]
a: [3,850,0,4,1,846]
s: 850
d: [2,680,1,2,1,1,40,256,415]
a: [3,688,1,2,1,680]
s: 688
d: [2,852,0,4,1,2,22,435,385]
a: [3,856,0,4,1,852]
s: 856
i: [1,274,0,5,2,0,0,2] t6_boole
s: 274
d: [2,682,1,2,1,1,40,256,462]
a: [3,698,1,2,1,682]
s: 698
i: [1,275,0,5,2,0,0,3] cc8_ordinal1
s: 275
d: [2,705,1,2,1,1,40,257,415]
a: [3,713,1,2,1,705]
s: 713
i: [1,276,0,5,2,0,0,2] cc1_xreal_0
s: 276
i: [1,277,0,5,2,0,0,2] cc1_xcmplx_0
s: 277
d: [2,707,1,2,1,1,40,257,462]
a: [3,723,1,2,1,707]
s: 723
i: [1,298,0,5,2,0,0,255] t7_boole
s: 298
d: [2,732,1,2,1,1,40,259,467]
a: [3,745,1,2,1,732]
s: 745
i: [1,244,0,5,1,0,0,3] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,410,0,5,1,244]
s: 410
i: [1,316,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,447,0,5,2,316]
s: 447
d: [2,749,1,2,1,1,40,260,467]
a: [3,762,1,2,1,749]
s: 762
d: [2,858,0,5,2,1,22,274,380]
s: 858
d: [2,901,1,7,2,2,44,355,858]
s: 901
i: [1,282,0,6,3,0,0,255] cc2_ordinal1
s: 282
i: [1,285,0,6,3,0,0,2] cc5_xxreal_0
s: 285
d: [2,902,2,8,2,3,40,901,447]
a: [3,906,2,4,1,902]
s: 906
i: [1,287,0,6,3,0,0,2] cc5_xxreal_0
s: 287
d: [2,764,1,2,1,1,40,271,618]
a: [3,771,1,2,1,764]
s: 771
i: [1,288,0,6,3,0,0,2] cc3_xxreal_0
s: 288
i: [1,290,0,6,3,0,0,2] cc3_xxreal_0
s: 290
d: [2,765,1,2,1,1,40,271,482]
a: [3,776,1,2,1,765]
s: 776
i: [1,292,0,6,3,0,0,2] cc7_xxreal_0
s: 292
d: [2,766,1,2,1,1,40,271,602]
a: [3,781,1,2,1,766]
s: 781
i: [1,293,0,6,3,0,0,2] cc7_xxreal_0
s: 293
i: [1,294,0,6,2,0,0,255] t1_subset
s: 294
d: [2,767,1,2,1,1,40,271,612]
a: [3,786,1,2,1,767]
s: 786
i: [1,303,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 303
d: [2,788,1,2,1,1,40,272,618]
a: [3,795,1,2,1,788]
s: 795
i: [1,258,0,7,3,0,0,4] cc5_ordinal1
s: 258
i: [1,269,0,7,2,0,0,1] t2_arithm
s: 269
d: [2,789,1,2,1,1,40,272,482]
a: [3,800,1,2,1,789]
s: 800
i: [1,270,0,7,2,0,0,3] t3_arithm
s: 270
d: [2,790,1,2,1,1,40,272,602]
a: [3,805,1,2,1,790]
s: 805
i: [1,296,0,7,2,0,0,255] t3_subset
s: 296
i: [1,297,0,7,3,0,0,2] t8_boole
s: 297
d: [2,791,1,2,1,1,40,272,612]
a: [3,810,1,2,1,791]
s: 810
i: [1,251,0,8,3,0,0,255] cc10_ordinal1
s: 251
d: [2,816,1,2,1,1,40,273,618]
a: [3,823,1,2,1,816]
s: 823
i: [1,279,0,8,4,0,0,2] cc4_xxreal_0
s: 279
i: [1,281,0,8,4,0,0,2] cc6_xxreal_0
s: 281
d: [2,1065,2,6,3,1,40,279,756]
a: [3,1079,2,2,1,1065]
s: 1079
r: 1079
i: [1,283,0,8,4,0,0,2] cc8_xxreal_0
s: 283
a: [3,1078,2,2,1,1065]
s: 1078
i: [1,295,0,8,3,0,0,255] t2_subset
s: 295
i: [1,313,0,8,3,0,0,1] fc6_xreal_0
s: 313
a: [3,1083,2,2,1,1065]
s: 1083
i: [1,314,0,8,3,0,0,1] fc3_xcmplx_0
s: 314
d: [2,817,1,2,1,1,40,273,482]
a: [3,828,1,2,1,817]
s: 828
i: [1,300,0,9,3,0,0,255] t5_subset
s: 300
i: [1,299,0,10,3,0,0,255] t4_subset
s: 299
d: [2,818,1,2,1,1,40,273,602]
a: [3,833,1,2,1,818]
s: 833
i: [1,317,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 317
d: [2,819,1,2,1,1,40,273,612]
a: [3,838,1,2,1,819]
s: 838
i: [1,263,0,11,5,0,0,1] t2_real
s: 263
i: [1,264,0,11,5,0,0,1] t1_real
s: 264
d: [2,860,1,2,1,1,40,275,630]
a: [3,866,1,2,1,860]
s: 866
i: [1,265,0,11,5,0,0,1] t8_real
s: 265
d: [2,1274,1,12,4,1,40,265,355]
a: [3,1283,1,4,1,1274]
s: 1283
d: [2,1297,2,0,0,2,52,360,360,1283,313]
e: 1297
d: [2,1298,2,4,2,1,40,1283,313]
d: [2,1299,2,2,1,1,18,1298]
d: [2,1301,2,0,0,2,21,1299,360]
e: 1301
r: 1283
i: [1,266,0,11,5,0,0,1] t7_real
s: 266
i: [1,267,0,11,5,0,0,1] t3_real
s: 267
a: [3,1282,1,4,1,1274]
s: 1282
d: [2,1303,3,0,0,2,52,906,1282,260]
e: 1303
d: [2,1304,2,4,1,1,40,1282,260]
d: [2,1307,2,0,0,2,21,1304,906]
e: 1307
r: 906
i: [1,268,0,11,5,0,0,1] t4_real
s: 268
a: [3,1291,1,4,1,1274]
s: 1291
i: [1,315,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 315
i: [1,304,0,12,5,0,0,1] fc26_xreal_0
s: 304
a: [3,1087,2,2,1,1065]
s: 1087
d: [2,1378,3,0,0,3,52,756,1083,1087,287]
e: 1378
d: [2,1379,3,4,2,1,40,1087,287]
d: [2,1382,3,2,1,2,21,1379,1083]
d: [2,1383,3,0,0,3,21,1382,756]
e: 1383
r: 1083
i: [1,305,0,12,5,0,0,1] fc24_xreal_0
s: 305
a: [3,1082,2,2,1,1065]
s: 1082
i: [1,306,0,12,5,0,0,1] fc23_xreal_0
s: 306
i: [1,307,0,12,5,0,0,1] fc25_xreal_0
s: 307
d: [2,1412,3,0,0,3,52,1082,360,1082,360,1291,307]
e: 1412
d: [2,1413,2,8,4,1,40,307,1291]
d: [2,1420,2,4,2,1,18,1413]
d: [2,1422,2,2,1,2,21,1420,360]
d: [2,1423,2,0,0,3,21,1422,1082]
e: 1423
r: 1082
r: 1087
r: 1291
a: [3,1287,1,2,1,1274]
s: 1287
d: [2,859,1,2,1,1,40,275,624]
d: [2,959,2,2,1,1,40,795,255]
d: [2,961,2,2,1,1,40,795,253]
d: [2,1368,2,8,4,1,40,304,1291]
d: [2,1375,2,4,2,1,18,1368]
d: [2,1376,2,2,1,2,21,1375,360]
d: [2,1385,3,2,1,2,21,1379,756]
d: [2,1427,2,4,2,1,40,1287,288]
d: [2,1429,2,2,1,2,21,1427,618]
f: 1429,1423,1385,1376,1301,1274,961,959,859,747,620,614,241,238,237,233,234
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t63_xreal_1
% SZS output start Proof for xreal_1__t63_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X0))) [input t63_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X0))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
14. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
19. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
41. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
44. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v3_xxreal_0(k3_xcmplx_0(X0,X1))) [input fc25_xreal_0]
45. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k3_xcmplx_0(X0,X1))) [input fc26_xreal_0]
48. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k3_xcmplx_0(X0,X1))) [input fc6_xreal_0]
68. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
77. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
95. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
107. ? [X0] : (~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X0)) & v1_xreal_0(X0)) [ennf transformation 2]
110. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
112. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
116. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
125. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 95]
126. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 125]
139. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
150. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 19]
151. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 150]
152. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 14]
153. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 152]
166. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 45]
167. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 166]
172. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 44]
173. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 172]
182. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 48]
183. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 182]
194. ? [X0] : (~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X0)) & v1_xreal_0(X0)) => (~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK0)) & v1_xreal_0(sK0)) [choice axiom]
195. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK0)) & v1_xreal_0(sK0) [skolemisation 107,194]
233. v1_xreal_0(sK0) [cnf transformation 195] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
234. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK0)) [cnf transformation 195] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
237. v1_xboole_0(k1_xboole_0) [cnf transformation 41] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
238. k6_numbers = k1_xboole_0 [cnf transformation 68] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
241. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 77] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
253. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 110] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
255. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 112] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
260. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 116] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
265. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 126] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
275. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 139] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
287. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 151] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
288. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 153] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
304. ~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 167] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
307. ~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 173] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
313. v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 183] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
353. 1 <=> r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK0)) [avatar definition]
355. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK0)) <- (~1) [avatar component clause 353] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
356. ~1 [avatar split clause 234,353]
358. 2 <=> v1_xreal_0(sK0) [avatar definition]
360. v1_xreal_0(sK0) <- (2) [avatar component clause 358] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
361. 2 [avatar split clause 233,358]
373. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
375. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 373] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
376. 5 [avatar split clause 237,373]
378. 6 <=> k6_numbers = k1_xboole_0 [avatar definition]
380. k6_numbers = k1_xboole_0 <- (6) [avatar component clause 378] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
381. 6 [avatar split clause 238,378]
393. 9 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
395. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (9) [avatar component clause 393] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
396. 9 [avatar split clause 241,393]
614. v1_xboole_0(k6_numbers) <- (5, 6) [backward demodulation 375,380] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
616. 54 <=> v1_xboole_0(k6_numbers) [avatar definition]
618. v1_xboole_0(k6_numbers) <- (54) [avatar component clause 616] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
619. 54 | ~5 | ~6 [avatar split clause 614,378,373,616]
620. m1_subset_1(k6_numbers,k4_ordinal1) <- (6, 9) [forward demodulation 395,380] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
622. 55 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
624. m1_subset_1(k6_numbers,k4_ordinal1) <- (55) [avatar component clause 622] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
625. 55 | ~6 | ~9 [avatar split clause 620,393,378,622]
747. v1_xxreal_0(sK0) <- (2) [resolution 260,360] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
754. 75 <=> v1_xxreal_0(sK0) [avatar definition]
756. v1_xxreal_0(sK0) <- (75) [avatar component clause 754] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
757. 75 | ~2 [avatar split clause 747,358,754]
793. 81 <=> v7_ordinal1(k6_numbers) [avatar definition]
795. v7_ordinal1(k6_numbers) <- (81) [avatar component clause 793] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
859. v7_ordinal1(k6_numbers) <- (55) [resolution 275,624] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
862. 81 | ~55 [avatar split clause 859,622,793]
959. v1_xxreal_0(k6_numbers) <- (81) [resolution 795,255] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
961. v1_xreal_0(k6_numbers) <- (81) [resolution 795,253] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
965. 106 <=> v1_xxreal_0(k6_numbers) [avatar definition]
968. 106 | ~81 [avatar split clause 959,793,965]
975. 108 <=> v1_xreal_0(k6_numbers) [avatar definition]
978. 108 | ~81 [avatar split clause 961,793,975]
1081. 121 <=> v2_xxreal_0(sK0) [avatar definition]
1082. ~v2_xxreal_0(sK0) <- (~121) [avatar component clause 1081] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1085. 122 <=> v3_xxreal_0(sK0) [avatar definition]
1087. v3_xxreal_0(sK0) <- (122) [avatar component clause 1085] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1274. v3_xxreal_0(k3_xcmplx_0(sK0,sK0)) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k3_xcmplx_0(sK0,sK0)) | ~v1_xreal_0(k6_numbers) <- (~1) [resolution 265,355] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1281. 149 <=> v1_xreal_0(k3_xcmplx_0(sK0,sK0)) [avatar definition]
1283. ~v1_xreal_0(k3_xcmplx_0(sK0,sK0)) <- (~149) [avatar component clause 1281] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1285. 150 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1287. v2_xxreal_0(k6_numbers) <- (150) [avatar component clause 1285] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1289. 151 <=> v3_xxreal_0(k3_xcmplx_0(sK0,sK0)) [avatar definition]
1291. v3_xxreal_0(k3_xcmplx_0(sK0,sK0)) <- (151) [avatar component clause 1289] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1292. ~108 | ~149 | 150 | 151 | 1 [avatar split clause 1274,353,1289,1285,1281,975]
1298. ~v1_xreal_0(sK0) | ~v1_xreal_0(sK0) <- (~149) [resolution 1283,313] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1299. ~v1_xreal_0(sK0) <- (~149) [duplicate literal removal 1298] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1301. $false <- (2, ~149) [subsumption resolution 1299,360] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1302. ~2 | 149 [avatar contradiction clause 1301]
1368. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (151) [resolution 304,1291] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1375. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (151) [duplicate literal removal 1368] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1376. v3_xxreal_0(sK0) <- (2, 151) [subsumption resolution 1375,360] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1377. 122 | ~2 | ~151 [avatar split clause 1376,1289,358,1085]
1379. ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (122) [resolution 1087,287] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1385. ~v2_xxreal_0(sK0) <- (75, 122) [subsumption resolution 1379,756] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1386. ~121 | ~75 | ~122 [avatar split clause 1385,1085,754,1081]
1413. ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) | ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (151) [resolution 307,1291] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1420. ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (151) [duplicate literal removal 1413] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1422. v2_xxreal_0(sK0) <- (2, 151) [subsumption resolution 1420,360] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1423. $false <- (2, ~121, 151) [subsumption resolution 1422,1082] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1424. ~2 | 121 | ~151 [avatar contradiction clause 1423]
1427. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (150) [resolution 1287,288] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1429. ~v1_xxreal_0(k6_numbers) <- (54, 150) [subsumption resolution 1427,618] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1430. ~106 | ~54 | ~150 [avatar split clause 1429,1285,616,965]
1431. $false [avatar sat refutation 356,361,376,381,396,619,625,757,862,968,978,1292,1302,1377,1386,1424,1430] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:839,thDist:-839}
% SZS output end Proof for xreal_1__t63_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 94
% >>> Preprocessing
% Introduced skolems: 19
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 119
% Generated clauses: 547
% Active clauses: 138
% Passive clauses: 238
% Final active clauses: 130
% Final passive clauses: 85
% Main loop iterations started: 152
% 
% >>> Simplifying Inferences
% Duplicate literals: 36
% Fw subsumption resolutions: 31
% Fw demodulations: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 26
% Forward subsumptions: 53
% 
% >>> Generating Inferences
% Binary resolution: 150
% Unit resulting resolution: 4
% Factoring: 1
% Forward superposition: 35
% Backward superposition: 12
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 26
% Split components: 64
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 179
% SAT solver unit clauses: 54
% SAT solver binary clauses: 82
% 
% Memory used [KB]: 5756
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% sat proof minimization: 0.001 s
% splitting component index usage: 0.001 s
% resolution: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

