i: [1,241,0,2,1,1,0,0] l206_xreal_1
a: [3,387,0,2,1,241]
s: 387
i: [1,240,0,2,1,1,0,0] l206_xreal_1
a: [3,392,0,2,1,240]
s: 392
i: [1,243,0,3,1,1,0,0] l206_xreal_1
a: [3,377,0,3,1,243]
s: 377
i: [1,242,0,3,1,1,0,0] l206_xreal_1
a: [3,382,0,3,1,242]
s: 382
i: [1,244,0,5,1,1,0,0] l206_xreal_1
a: [3,372,0,5,1,244]
s: 372
i: [1,245,0,2,1,0,0,255] fc1_numbers
a: [3,397,0,2,1,245]
s: 397
i: [1,246,0,2,1,0,0,1] spc1_boole
a: [3,402,0,2,1,246]
s: 402
i: [1,247,0,2,1,0,0,255] fc1_xboole_0
a: [3,407,0,2,1,247]
s: 407
i: [1,248,0,2,1,0,0,255] spc0_boole
a: [3,412,0,2,1,248]
s: 412
i: [1,260,0,2,1,0,0,255] fc6_ordinal1
a: [3,467,0,2,1,260]
s: 467
i: [1,259,0,2,1,0,0,255] fc6_ordinal1
a: [3,472,0,2,1,259]
s: 472
i: [1,261,0,2,1,0,0,255] spc1_numerals
a: [3,492,0,2,1,261]
s: 492
i: [1,336,0,2,1,0,0,2] rc2_xboole_0
a: [3,519,0,2,1,336]
s: 519
i: [1,337,0,2,1,0,0,255] rc4_ordinal1
a: [3,524,0,2,1,337]
s: 524
i: [1,338,0,2,1,0,0,255] rc1_ordinal1
a: [3,529,0,2,1,338]
s: 529
i: [1,339,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,534,0,2,1,339]
s: 534
i: [1,340,0,2,1,0,0,1] rc1_xreal_0
a: [3,539,0,2,1,340]
s: 539
i: [1,341,0,2,1,0,0,2] rc1_xxreal_0
a: [3,544,0,2,1,341]
s: 544
i: [1,342,0,2,1,0,0,2] rc1_xboole_0
a: [3,549,0,2,1,342]
s: 549
i: [1,346,0,2,1,0,0,255] rc2_ordinal1
a: [3,554,0,2,1,346]
s: 554
i: [1,345,0,2,1,0,0,255] rc2_ordinal1
a: [3,559,0,2,1,345]
s: 559
i: [1,344,0,2,1,0,0,255] rc2_ordinal1
a: [3,564,0,2,1,344]
s: 564
i: [1,343,0,2,1,0,0,255] rc2_ordinal1
a: [3,569,0,2,1,343]
s: 569
i: [1,348,0,2,1,0,0,255] rc5_ordinal1
a: [3,574,0,2,1,348]
s: 574
i: [1,347,0,2,1,0,0,255] rc5_ordinal1
a: [3,579,0,2,1,347]
s: 579
i: [1,350,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,584,0,2,1,350]
s: 584
i: [1,349,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,589,0,2,1,349]
s: 589
i: [1,354,0,2,1,0,0,2] rc3_xreal_0
a: [3,594,0,2,1,354]
s: 594
i: [1,353,0,2,1,0,0,2] rc3_xreal_0
a: [3,599,0,2,1,353]
s: 599
i: [1,352,0,2,1,0,0,2] rc3_xreal_0
a: [3,604,0,2,1,352]
s: 604
i: [1,351,0,2,1,0,0,2] rc3_xreal_0
a: [3,609,0,2,1,351]
s: 609
i: [1,358,0,2,1,0,0,2] rc2_xreal_0
a: [3,614,0,2,1,358]
s: 614
i: [1,357,0,2,1,0,0,2] rc2_xreal_0
a: [3,619,0,2,1,357]
s: 619
i: [1,356,0,2,1,0,0,2] rc2_xreal_0
a: [3,624,0,2,1,356]
s: 624
i: [1,355,0,2,1,0,0,2] rc2_xreal_0
a: [3,629,0,2,1,355]
s: 629
i: [1,360,0,2,1,0,0,2] rc3_xxreal_0
a: [3,634,0,2,1,360]
s: 634
i: [1,359,0,2,1,0,0,2] rc3_xxreal_0
a: [3,639,0,2,1,359]
s: 639
i: [1,362,0,2,1,0,0,2] rc2_xxreal_0
a: [3,644,0,2,1,362]
s: 644
i: [1,361,0,2,1,0,0,2] rc2_xxreal_0
a: [3,649,0,2,1,361]
s: 649
i: [1,366,0,2,1,0,0,2] rc4_xreal_0
a: [3,654,0,2,1,366]
s: 654
i: [1,365,0,2,1,0,0,2] rc4_xreal_0
a: [3,659,0,2,1,365]
s: 659
i: [1,364,0,2,1,0,0,2] rc4_xreal_0
a: [3,664,0,2,1,364]
s: 664
i: [1,363,0,2,1,0,0,2] rc4_xreal_0
a: [3,669,0,2,1,363]
s: 669
i: [1,368,0,2,1,0,0,2] rc4_xxreal_0
a: [3,674,0,2,1,368]
s: 674
i: [1,367,0,2,1,0,0,2] rc4_xxreal_0
a: [3,679,0,2,1,367]
s: 679
i: [1,321,0,3,1,0,0,255] reflexivity_r1_tarski
s: 321
i: [1,249,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r1_r0
a: [3,417,0,3,1,249]
s: 417
i: [1,250,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,422,0,3,1,250]
s: 422
i: [1,251,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,427,0,3,1,251]
s: 427
r: 407
d: [2,681,0,2,1,2,23,407,427]
a: [3,685,0,2,1,681]
s: 685
i: [1,252,0,3,1,0,0,255] t1_numerals
a: [3,432,0,3,1,252]
s: 432
i: [1,253,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r0_r0
a: [3,437,0,3,1,253]
s: 437
i: [1,254,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r0_r1
a: [3,442,0,3,1,254]
s: 442
i: [1,255,0,3,1,0,0,1] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,447,0,3,1,255]
s: 447
i: [1,264,0,3,1,0,0,255] spc1_numerals
a: [3,477,0,3,1,264]
s: 477
i: [1,263,0,3,1,0,0,255] spc1_numerals
a: [3,482,0,3,1,263]
s: 482
i: [1,267,0,3,1,0,0,255] spc0_numerals
a: [3,497,0,3,1,267]
s: 497
i: [1,266,0,3,1,0,0,255] spc0_numerals
a: [3,502,0,3,1,266]
s: 502
d: [2,687,0,3,1,2,22,432,427]
a: [3,691,0,3,1,687]
s: 691
d: [2,693,0,3,1,2,22,482,422]
a: [3,697,0,3,1,693]
s: 697
d: [2,699,0,3,1,2,22,502,422]
a: [3,703,0,3,1,699]
s: 703
i: [1,269,0,4,2,0,0,255] cc6_ordinal1
s: 269
i: [1,270,0,4,2,0,0,255] cc2_xcmplx_0
s: 270
d: [2,705,1,2,1,1,40,269,524]
a: [3,710,1,2,1,705]
s: 710
i: [1,271,0,4,2,0,0,255] cc2_xreal_0
s: 271
i: [1,272,0,4,2,0,0,255] cc2_xxreal_0
s: 272
d: [2,706,1,2,1,1,40,269,574]
a: [3,715,1,2,1,706]
s: 715
i: [1,273,0,4,2,0,0,255] cc1_ordinal1
s: 273
d: [2,717,1,2,1,1,40,270,524]
a: [3,722,1,2,1,717]
s: 722
i: [1,274,0,4,2,0,0,255] cc1_ordinal1
s: 274
i: [1,279,0,4,2,0,0,2] cc3_xreal_0
s: 279
d: [2,803,1,2,1,1,40,279,392]
a: [3,812,1,2,1,803]
s: 812
i: [1,280,0,4,2,0,0,2] cc4_xreal_0
s: 280
d: [2,804,1,2,1,1,40,279,387]
a: [3,817,1,2,1,804]
s: 817
i: [1,290,0,4,2,0,0,255] cc9_ordinal1
s: 290
i: [1,291,0,4,2,0,0,255] cc7_ordinal1
s: 291
d: [2,824,1,2,1,1,40,280,392]
a: [3,833,1,2,1,824]
s: 833
i: [1,292,0,4,2,0,0,255] cc3_ordinal1
s: 292
d: [2,825,1,2,1,1,40,280,387]
a: [3,838,1,2,1,825]
s: 838
i: [1,320,0,4,1,0,0,255] existence_m1_subset_1
s: 320
i: [1,256,0,4,1,0,0,255] dt_k5_numbers
a: [3,452,0,4,1,256]
s: 452
d: [2,718,1,2,1,1,40,270,574]
a: [3,727,1,2,1,718]
s: 727
i: [1,257,0,4,1,0,0,255] dt_k6_numbers
a: [3,457,0,4,1,257]
s: 457
d: [2,729,1,2,1,1,40,271,524]
a: [3,734,1,2,1,729]
s: 734
i: [1,262,0,4,1,0,0,255] spc1_numerals
a: [3,487,0,4,1,262]
s: 487
i: [1,265,0,4,1,0,0,255] spc0_numerals
a: [3,507,0,4,1,265]
s: 507
d: [2,730,1,2,1,1,40,271,574]
a: [3,739,1,2,1,730]
s: 739
d: [2,935,0,4,1,2,22,452,422]
a: [3,939,0,4,1,935]
s: 939
d: [2,741,1,2,1,1,40,272,524]
a: [3,746,1,2,1,741]
s: 746
d: [2,941,0,4,1,2,22,457,422]
a: [3,945,0,4,1,941]
s: 945
d: [2,950,0,4,1,2,22,487,422]
a: [3,954,0,4,1,950]
s: 954
d: [2,742,1,2,1,1,40,272,574]
a: [3,751,1,2,1,742]
s: 751
d: [2,956,0,4,1,2,22,507,422]
a: [3,960,0,4,1,956]
s: 960
d: [2,753,1,2,1,1,40,273,467]
a: [3,761,1,2,1,753]
s: 761
i: [1,293,0,5,2,0,0,255] t6_boole
s: 293
i: [1,294,0,5,2,0,0,255] cc8_ordinal1
s: 294
d: [2,755,1,2,1,1,40,273,529]
a: [3,771,1,2,1,755]
s: 771
i: [1,295,0,5,2,0,0,255] cc1_xcmplx_0
s: 295
d: [2,778,1,2,1,1,40,274,467]
a: [3,786,1,2,1,778]
s: 786
i: [1,296,0,5,2,0,0,255] cc1_xreal_0
s: 296
i: [1,317,0,5,2,0,0,255] t7_boole
s: 317
d: [2,780,1,2,1,1,40,274,529]
a: [3,796,1,2,1,780]
s: 796
i: [1,258,0,5,1,0,0,2] rqRealDiv__k7_xcmplx_0__r1_r1_r1
a: [3,462,0,5,1,258]
s: 462
d: [2,805,1,2,1,1,40,279,539]
a: [3,822,1,2,1,805]
s: 822
i: [1,334,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,514,0,5,2,334]
s: 514
d: [2,965,0,5,2,1,22,293,427]
s: 965
d: [2,1022,1,5,2,2,44,377,965]
s: 1022
i: [1,301,0,6,3,0,0,255] cc2_ordinal1
s: 301
d: [2,826,1,2,1,1,40,280,539]
a: [3,843,1,2,1,826]
s: 843
i: [1,304,0,6,3,0,0,2] cc5_xxreal_0
s: 304
i: [1,306,0,6,3,0,0,2] cc5_xxreal_0
s: 306
d: [2,845,1,2,1,1,40,290,685]
a: [3,853,1,2,1,845]
s: 853
i: [1,307,0,6,3,0,0,2] cc3_xxreal_0
s: 307
d: [2,846,1,2,1,1,40,290,412]
a: [3,858,1,2,1,846]
s: 858
i: [1,309,0,6,3,0,0,2] cc3_xxreal_0
s: 309
i: [1,311,0,6,3,0,0,2] cc7_xxreal_0
s: 311
d: [2,847,1,2,1,1,40,290,549]
a: [3,863,1,2,1,847]
s: 863
i: [1,312,0,6,3,0,0,2] cc7_xxreal_0
s: 312
d: [2,848,1,2,1,1,40,290,669]
a: [3,868,1,2,1,848]
s: 868
i: [1,313,0,6,2,0,0,255] t1_subset
s: 313
i: [1,322,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 322
d: [2,849,1,2,1,1,40,290,679]
a: [3,873,1,2,1,849]
s: 873
i: [1,275,0,7,3,0,0,255] cc5_ordinal1
s: 275
d: [2,875,1,2,1,1,40,291,685]
a: [3,883,1,2,1,875]
s: 883
i: [1,276,0,7,2,0,0,1] t5_arithm
s: 276
i: [1,277,0,7,2,0,0,2] t6_arithm
s: 277
d: [2,876,1,2,1,1,40,291,412]
a: [3,888,1,2,1,876]
s: 888
i: [1,315,0,7,2,0,0,255] t3_subset
s: 315
d: [2,877,1,2,1,1,40,291,549]
a: [3,893,1,2,1,877]
s: 893
i: [1,316,0,7,3,0,0,2] t8_boole
s: 316
i: [1,268,0,8,3,0,0,255] cc10_ordinal1
s: 268
d: [2,878,1,2,1,1,40,291,669]
a: [3,898,1,2,1,878]
s: 898
i: [1,298,0,8,4,0,0,2] cc4_xxreal_0
s: 298
d: [2,1165,2,6,3,1,40,298,833]
a: [3,1180,2,2,1,1165]
s: 1180
r: 1180
i: [1,300,0,8,4,0,0,2] cc6_xxreal_0
s: 300
i: [1,302,0,8,4,0,0,2] cc8_xxreal_0
s: 302
a: [3,1179,2,2,1,1165]
s: 1179
i: [1,314,0,8,3,0,0,255] t2_subset
s: 314
a: [3,1184,2,2,1,1165]
s: 1184
i: [1,332,0,8,3,0,0,1] fc5_xcmplx_0
s: 332
i: [1,333,0,8,3,0,0,1] fc8_xreal_0
s: 333
d: [2,1166,2,6,3,1,40,298,838]
a: [3,1193,2,2,1,1166]
s: 1193
r: 1193
i: [1,319,0,9,3,0,0,255] t5_subset
s: 319
a: [3,1192,2,2,1,1166]
s: 1192
i: [1,278,0,10,3,0,0,1] t60_xcmplx_1
s: 278
i: [1,318,0,10,3,0,0,255] t4_subset
s: 318
a: [3,1197,2,2,1,1166]
s: 1197
i: [1,335,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 335
d: [2,1370,1,11,3,1,40,335,372]
a: [3,1377,1,2,1,1370]
s: 1377
i: [1,283,0,11,5,0,0,1] t2_real
s: 283
i: [1,284,0,11,5,0,0,1] t1_real
s: 284
a: [3,1187,2,2,1,1165]
s: 1187
i: [1,285,0,11,5,0,0,1] t8_real
s: 285
a: [3,1200,2,2,1,1166]
s: 1200
i: [1,286,0,11,5,0,0,1] t7_real
s: 286
i: [1,287,0,11,5,0,0,1] t3_real
s: 287
d: [2,1436,1,12,4,1,40,285,372]
a: [3,1445,1,4,1,1436]
s: 1445
d: [2,1451,2,0,0,3,52,392,387,1445,333]
e: 1451
d: [2,1452,2,4,2,1,40,1445,333]
d: [2,1454,2,2,1,2,21,1452,387]
d: [2,1455,2,0,0,3,21,1454,392]
e: 1455
r: 1445
i: [1,288,0,11,5,0,0,1] t4_real
s: 288
a: [3,1444,1,4,1,1436]
s: 1444
i: [1,323,0,12,5,0,0,1] fc33_xreal_0
s: 323
i: [1,324,0,12,5,0,0,1] fc31_xreal_0
s: 324
a: [3,1380,1,4,1,1370]
s: 1380
i: [1,325,0,12,5,0,0,1] fc32_xreal_0
s: 325
a: [3,1449,1,4,1,1436]
s: 1449
i: [1,326,0,12,5,0,0,1] fc34_xreal_0
s: 326
i: [1,331,0,12,5,0,0,1] fc9_xcmplx_0
s: 331
d: [2,1390,1,8,4,1,40,283,377]
d: [2,1405,1,6,3,2,21,1390,392]
a: [3,1408,1,2,1,1405]
s: 1408
i: [1,281,0,13,6,0,0,1] t5_real
s: 281
d: [2,1458,2,4,1,1,40,1444,279]
a: [3,1463,2,4,1,1458]
s: 1463
i: [1,282,0,13,6,0,0,1] t6_real
s: 282
i: [1,330,0,14,4,0,0,255] existence_m2_subset_1
s: 330
d: [2,1493,2,12,3,1,40,1380,298]
a: [3,1496,2,4,1,1493]
s: 1496
i: [1,327,0,15,5,0,0,255] dt_m2_subset_1
s: 327
d: [2,1537,2,4,2,2,44,1408,965]
s: 1537
r: 304
i: [1,328,0,15,5,0,0,255] redefinition_m2_subset_1
s: 328
i: [1,329,0,15,5,0,0,255] redefinition_m2_subset_1
s: 329
d: [2,879,1,2,1,1,40,291,679]
a: [3,903,1,2,1,879]
s: 903
i: [1,289,0,19,6,0,0,1] l66_xreal_1
s: 289
d: [2,905,1,2,1,1,40,292,685]
a: [3,913,1,2,1,905]
s: 913
d: [2,906,1,2,1,1,40,292,412]
a: [3,918,1,2,1,906]
s: 918
d: [2,907,1,2,1,1,40,292,549]
a: [3,923,1,2,1,907]
s: 923
d: [2,908,1,2,1,1,40,292,669]
a: [3,928,1,2,1,908]
s: 928
d: [2,909,1,2,1,1,40,292,679]
a: [3,933,1,2,1,909]
s: 933
d: [2,967,1,2,1,1,40,294,697]
a: [3,974,1,2,1,967]
s: 974
d: [2,1733,2,0,0,2,52,1377,974,272]
e: 1733
d: [2,1734,2,2,1,1,40,974,272]
d: [2,1739,2,0,0,2,21,1734,1377]
e: 1739
r: 1377
a: [3,1376,1,2,1,1370]
s: 1376
a: [3,1385,1,5,1,1370]
s: 1385
a: [3,1500,2,4,1,1493]
s: 1500
d: [2,982,1,2,1,1,40,295,477]
a: [3,988,1,2,1,982]
s: 988
d: [2,983,1,2,1,1,40,295,497]
a: [3,993,1,2,1,983]
s: 993
d: [2,1000,1,2,1,1,40,296,477]
a: [3,1006,1,2,1,1000]
s: 1006
d: [2,1001,1,2,1,1,40,296,497]
a: [3,1011,1,2,1,1001]
s: 1011
d: [2,1026,1,4,2,1,40,304,599]
d: [2,1028,1,2,1,2,21,1026,604]
a: [3,1032,1,2,1,1028]
s: 1032
d: [2,1027,1,4,2,1,40,304,634]
d: [2,1034,1,2,1,2,21,1027,639]
a: [3,1038,1,2,1,1034]
s: 1038
d: [2,1040,1,4,2,1,40,306,599]
d: [2,1042,1,2,1,2,21,1040,604]
a: [3,1046,1,2,1,1042]
s: 1046
d: [2,1041,1,4,2,1,40,306,634]
d: [2,1048,1,2,1,2,21,1041,639]
a: [3,1052,1,2,1,1048]
s: 1052
d: [2,1056,1,4,2,1,40,307,619]
d: [2,1058,1,2,1,2,21,1056,624]
a: [3,1062,1,2,1,1058]
s: 1062
d: [2,1057,1,4,2,1,40,307,644]
d: [2,1064,1,2,1,2,21,1057,649]
a: [3,1068,1,2,1,1064]
s: 1068
d: [2,1072,1,4,2,1,40,275,477]
a: [3,1084,1,2,1,1072]
s: 1084
d: [2,1148,1,4,2,1,40,268,939]
a: [3,1154,1,2,1,1148]
s: 1154
d: [2,1169,1,6,3,1,40,298,544]
a: [3,1232,1,2,1,1169]
s: 1232
r: 1232
a: [3,1231,1,2,1,1169]
s: 1231
a: [3,1236,1,2,1,1169]
s: 1236
d: [2,1393,1,8,4,1,40,283,442]
a: [3,1414,1,2,1,1393]
s: 1414
d: [2,1788,3,0,0,4,52,412,1414,1537]
e: 1788
d: [2,1789,3,2,1,3,40,1414,1537]
d: [2,1792,3,0,0,4,21,1789,412]
e: 1792
r: 1414
a: [3,1418,1,2,1,1393]
s: 1418
d: [2,1477,1,12,6,0,44,324,276]
d: [2,1482,1,10,5,0,21,1477,279]
a: [3,1489,1,2,1,1482]
s: 1489
d: [2,969,1,3,1,0,40,294,320]
a: [3,980,1,3,1,969]
s: 980
d: [2,754,2,2,1,1,40,273,710]
a: [3,766,2,2,1,754]
s: 766
d: [2,984,1,3,1,0,40,295,320]
a: [3,998,1,3,1,984]
s: 998
d: [2,757,2,2,1,1,40,273,715]
a: [3,776,2,2,1,757]
s: 776
d: [2,1002,1,3,1,0,40,296,320]
a: [3,1016,1,3,1,1002]
s: 1016
d: [2,1273,1,5,2,1,40,314,477]
d: [2,1287,1,3,1,2,21,1273,397]
a: [3,1291,1,3,1,1287]
s: 1291
d: [2,779,2,2,1,1,40,274,710]
a: [3,791,2,2,1,779]
s: 791
d: [2,1274,1,5,2,1,40,314,497]
d: [2,1293,1,3,1,2,21,1274,397]
a: [3,1297,1,3,1,1293]
s: 1297
d: [2,782,2,2,1,1,40,274,715]
a: [3,801,2,2,1,782]
s: 801
d: [2,1275,1,5,2,1,40,314,691]
d: [2,1299,1,3,1,2,21,1275,472]
a: [3,1303,1,3,1,1299]
s: 1303
d: [2,1276,1,5,2,1,40,314,697]
d: [2,1305,1,3,1,2,21,1276,472]
a: [3,1309,1,3,1,1305]
s: 1309
d: [2,1093,2,2,1,1,40,883,272]
a: [3,1101,2,2,1,1093]
s: 1101
d: [2,1277,1,5,2,1,40,314,703]
d: [2,1311,1,3,1,2,21,1277,472]
a: [3,1315,1,3,1,1311]
s: 1315
d: [2,1094,2,2,1,1,40,883,271]
a: [3,1106,2,2,1,1094]
s: 1106
d: [2,1572,1,11,4,1,40,327,945]
d: [2,1577,1,7,3,2,21,1572,939]
d: [2,1578,1,5,2,3,21,1577,472]
d: [2,1579,1,3,1,4,21,1578,397]
a: [3,1583,1,3,1,1579]
s: 1583
d: [2,1618,1,17,6,1,44,289,462]
d: [2,1629,1,15,5,1,18,1618]
a: [3,1641,1,3,1,1629]
s: 1641
d: [2,1887,3,0,0,5,52,1106,1006,1489,1418,1641,285]
e: 1887
d: [2,1888,2,8,4,1,40,1641,285]
d: [2,1892,2,6,3,2,21,1888,1418]
d: [2,1893,2,4,2,3,21,1892,1489]
d: [2,1894,2,2,1,4,21,1893,1006]
d: [2,1895,2,0,0,5,21,1894,1106]
e: 1895
r: 1537
r: 1408
r: 1489
r: 1641
a: [3,1409,1,2,1,1405]
s: 1409
r: 1409
s: 304
s: 1408
a: [3,1488,1,2,1,1482]
s: 1488
d: [2,1914,3,0,0,3,52,1101,685,1488,307]
e: 1914
d: [2,1915,2,4,2,1,40,1488,307]
d: [2,1918,2,2,1,2,21,1915,685]
d: [2,1919,2,0,0,3,21,1918,1101]
e: 1919
r: 1488
s: 1489
d: [2,1913,2,4,2,2,44,1408,965]
s: 1913
r: 304
a: [3,1640,1,3,1,1629]
s: 1640
d: [2,1095,2,2,1,1,40,883,270]
a: [3,1111,2,2,1,1095]
s: 1111
d: [2,1271,1,7,2,1,40,314,939]
a: [3,1281,1,4,1,1271]
s: 1281
a: [3,1284,1,3,1,1271]
s: 1284
d: [2,1115,2,2,1,1,40,888,272]
a: [3,1122,2,2,1,1115]
s: 1122
a: [3,1485,1,4,2,1482]
s: 1485
r: 326
r: 325
r: 324
r: 285
r: 284
r: 282
r: 281
d: [2,1940,3,0,0,3,52,1489,1106,1485]
e: 1940
d: [2,1941,2,0,0,3,52,1046,594,1485]
e: 1941
d: [2,1942,3,2,1,2,40,1485,1106]
d: [2,1958,3,0,0,3,21,1942,1489]
e: 1958
d: [2,1953,2,2,1,2,40,1485,594]
d: [2,1960,2,0,0,3,21,1953,1046]
e: 1960
r: 1485
a: [3,1413,1,2,1,1393]
s: 1413
s: 284
s: 285
d: [2,1128,2,2,1,1,40,893,272]
a: [3,1135,2,2,1,1128]
s: 1135
s: 324
d: [2,1129,2,2,1,1,40,893,271]
a: [3,1140,2,2,1,1129]
s: 1140
s: 325
s: 326
d: [2,1130,2,2,1,1,40,893,270]
a: [3,1145,2,2,1,1130]
s: 1145
s: 281
d: [2,1167,2,6,3,1,40,298,746]
a: [3,1206,2,2,1,1167]
s: 1206
r: 1206
s: 282
d: [2,1020,1,5,2,2,44,691,965]
s: 1020
a: [3,1205,2,2,1,1167]
s: 1205
d: [2,1070,1,5,2,0,40,275,320]
s: 1070
a: [3,1210,2,2,1,1167]
s: 1210
d: [2,1019,1,6,2,2,44,945,965]
s: 1019
d: [2,1150,1,6,2,0,40,268,320]
s: 1150
d: [2,1168,2,6,3,1,40,298,843]
a: [3,1219,2,2,1,1168]
s: 1219
r: 1219
d: [2,1270,1,6,2,0,40,314,320]
s: 1270
a: [3,1218,2,2,1,1168]
s: 1218
d: [2,1366,1,6,2,1,40,318,939]
s: 1366
d: [2,1018,1,7,3,1,44,965,965]
s: 1018
a: [3,1223,2,2,1,1168]
s: 1223
d: [2,1149,1,7,3,0,40,268,315]
s: 1149
d: [2,1170,2,6,3,1,40,298,751]
d: [2,1242,2,4,2,2,21,1170,579]
a: [3,1246,2,2,1,1242]
s: 1246
d: [2,1350,1,7,2,0,40,319,320]
s: 1350
d: [2,1594,1,11,4,1,40,329,939]
d: [2,1597,1,9,3,2,21,1594,472]
d: [2,1598,1,7,2,3,21,1597,397]
s: 1598
d: [2,1600,2,2,1,1,40,903,271]
a: [3,1606,2,2,1,1600]
s: 1606
d: [2,1127,1,8,3,0,40,315,275]
s: 1127
d: [2,1601,2,2,1,1,40,903,270]
a: [3,1611,2,2,1,1601]
s: 1611
d: [2,1324,1,8,3,0,40,333,280]
s: 1324
d: [2,1325,1,8,3,0,40,333,279]
s: 1325
d: [2,1672,2,2,1,1,40,913,274]
a: [3,1678,2,2,1,1672]
s: 1678
d: [2,1349,1,8,3,0,40,319,315]
s: 1349
d: [2,1673,2,2,1,1,40,913,273]
a: [3,1683,2,2,1,1673]
s: 1683
d: [2,1368,1,8,2,0,40,318,320]
s: 1368
d: [2,1540,1,18,9,0,40,281,285]
d: [2,1549,1,14,7,0,18,1540]
a: [3,1553,1,8,4,1549]
s: 1553
r: 1553
d: [2,1685,2,2,1,1,40,918,274]
a: [3,1690,2,2,1,1685]
s: 1690
a: [3,1556,1,6,3,1549]
s: 1556
d: [2,2164,3,0,0,4,52,1408,1489,1106,1556]
e: 2164
d: [2,2165,3,4,2,2,40,1556,1106]
d: [2,2182,3,2,1,3,21,2165,1489]
d: [2,2183,3,0,0,4,21,2182,1408]
e: 2183
r: 1556
d: [2,1686,2,2,1,1,40,918,273]
a: [3,1695,2,2,1,1686]
s: 1695
s: 1553
d: [2,1114,1,9,3,1,44,276,965]
s: 1114
d: [2,2207,2,14,7,1,44,324,1114]
d: [2,2231,2,12,6,1,21,2207,279]
d: [2,2232,2,10,5,2,21,2231,1913]
a: [3,2235,2,6,3,2232]
s: 2235
d: [2,1367,1,9,3,0,40,318,315]
s: 1367
d: [2,1697,2,2,1,1,40,923,274]
a: [3,1702,2,2,1,1697]
s: 1702
d: [2,1479,1,15,7,0,44,324,278]
d: [2,1480,1,13,6,0,18,1479]
d: [2,1491,1,11,5,1,21,1480,492]
d: [2,1492,1,9,4,1,21,1491,279]
s: 1492
d: [2,1506,1,15,7,0,44,325,278]
d: [2,1507,1,13,6,0,18,1506]
d: [2,1509,1,11,5,1,21,1507,492]
d: [2,1510,1,9,4,1,21,1509,279]
s: 1510
d: [2,1698,2,2,1,1,40,923,273]
a: [3,1707,2,2,1,1698]
s: 1707
d: [2,1991,1,15,7,0,44,324,278]
d: [2,1992,1,13,6,0,18,1991]
d: [2,1994,1,11,5,1,21,1992,492]
d: [2,1995,1,9,4,1,21,1994,279]
s: 1995
d: [2,1709,2,2,1,1,40,928,274]
a: [3,1714,2,2,1,1709]
s: 1714
d: [2,2003,1,15,7,0,44,325,278]
d: [2,2004,1,13,6,0,18,2003]
d: [2,2006,1,11,5,1,21,2004,492]
d: [2,2007,1,9,4,1,21,2006,279]
s: 2007
d: [2,1272,1,10,3,0,40,314,315]
s: 1272
d: [2,1710,2,2,1,1,40,928,273]
a: [3,1719,2,2,1,1710]
s: 1719
a: [3,1644,1,10,3,1629]
s: 1644
d: [2,1721,2,2,1,1,40,933,274]
a: [3,1726,2,2,1,1721]
s: 1726
d: [2,1622,1,17,6,1,44,289,462]
d: [2,1627,1,15,5,1,18,1622]
a: [3,1659,1,10,3,1627]
s: 1659
d: [2,1389,1,15,7,0,40,283,335]
d: [2,1403,1,13,6,0,21,1389,280]
d: [2,1404,1,11,5,0,21,1403,280]
s: 1404
d: [2,1722,2,2,1,1,40,933,273]
a: [3,1731,2,2,1,1722]
s: 1731
d: [2,1425,1,15,7,0,40,284,335]
d: [2,1434,1,13,6,0,21,1425,280]
d: [2,1435,1,11,5,0,21,1434,280]
s: 1435
d: [2,1757,2,2,1,1,40,1084,274]
a: [3,1762,2,2,1,1757]
s: 1762
d: [2,1964,1,15,7,0,40,284,335]
d: [2,1976,1,13,6,0,21,1964,280]
d: [2,1977,1,11,5,0,21,1976,280]
s: 1977
d: [2,1620,1,19,7,0,44,289,277]
d: [2,1651,1,17,6,0,21,1620,279]
a: [3,1654,1,12,4,1651]
s: 1654
d: [2,1758,2,2,1,1,40,1084,273]
a: [3,1767,2,2,1,1758]
s: 1767
d: [2,1624,1,19,7,0,44,289,277]
d: [2,1666,1,17,6,0,21,1624,279]
a: [3,1669,1,12,4,1666]
s: 1669
d: [2,2252,2,9,4,2,44,377,1492]
s: 2252
d: [2,1575,1,21,7,0,40,327,330]
d: [2,1576,1,13,4,0,18,1575]
s: 1576
d: [2,1592,1,21,7,0,40,328,330]
d: [2,1593,1,13,4,0,18,1592]
s: 1593
d: [2,1795,2,3,1,1,40,980,272]
a: [3,1802,2,3,1,1795]
s: 1802
d: [2,1595,1,14,5,0,40,329,315]
s: 1595
d: [2,1796,2,3,1,1,40,980,271]
a: [3,1807,2,3,1,1796]
s: 1807
d: [2,1619,1,19,7,0,44,289,276]
d: [2,1646,1,17,6,0,21,1619,279]
a: [3,1649,1,15,5,1646]
s: 1649
d: [2,1623,1,19,7,0,44,289,276]
d: [2,1661,1,17,6,0,21,1623,279]
a: [3,1664,1,15,5,1661]
s: 1664
d: [2,1889,2,7,3,1,40,1641,335]
d: [2,1897,2,5,2,2,21,1889,1101]
d: [2,1898,2,3,1,3,21,1897,1376]
a: [3,1905,2,3,1,1898]
s: 1905
d: [2,2455,3,0,0,6,52,1006,1106,1418,402,1489,1905,282]
e: 2455
d: [2,2456,3,0,0,5,52,1006,1106,1489,492,1905,284]
e: 2456
d: [2,2457,3,10,5,1,40,1905,282]
d: [2,2465,3,8,4,2,21,2457,1489]
d: [2,2466,3,6,3,3,21,2465,402]
d: [2,2467,3,4,2,4,21,2466,1418]
d: [2,2468,3,2,1,5,21,2467,1106]
d: [2,2469,3,0,0,6,21,2468,1006]
e: 2469
d: [2,2459,3,8,4,1,40,1905,284]
d: [2,2471,3,6,3,2,21,2459,492]
d: [2,2472,3,4,2,3,21,2471,1489]
d: [2,2473,3,2,1,4,21,2472,1106]
d: [2,2474,3,0,0,5,21,2473,1006]
e: 2474
r: 1905
d: [2,1596,1,17,4,0,40,329,320]
s: 1596
a: [3,1904,2,3,1,1898]
s: 1904
d: [2,1621,1,22,8,0,44,289,278]
d: [2,1628,1,20,7,0,18,1621]
d: [2,1656,1,18,6,0,21,1628,279]
s: 1656
d: [2,1625,1,22,8,0,44,289,278]
d: [2,1626,1,20,7,0,18,1625]
d: [2,1671,1,18,6,0,21,1626,279]
s: 1671
d: [2,2516,2,13,5,1,40,1671,372]
d: [2,2529,2,10,4,2,21,2516,382]
d: [2,2530,2,8,3,3,21,2529,387]
d: [2,2531,2,6,2,4,21,2530,392]
a: [3,2539,2,3,1,2531]
s: 2539
d: [2,2542,3,0,0,5,52,1197,387,1106,1489,2539,1435]
e: 2542
d: [2,2543,3,0,0,5,52,1106,387,1489,1200,2539,285]
e: 2543
d: [2,2544,3,8,4,1,40,2539,1435]
d: [2,2552,3,6,3,2,21,2544,1489]
d: [2,2553,3,4,2,3,21,2552,1106]
d: [2,2554,3,2,1,4,21,2553,387]
d: [2,2555,3,0,0,5,21,2554,1197]
e: 2555
d: [2,2546,3,8,4,1,40,2539,285]
d: [2,2557,3,6,3,2,21,2546,1200]
d: [2,2558,3,4,2,3,21,2557,1489]
d: [2,2559,3,2,1,4,21,2558,387]
d: [2,2560,3,0,0,5,21,2559,1106]
e: 2560
r: 1192
r: 1197
r: 2539
d: [2,966,1,2,1,1,40,294,691]
d: [2,1394,1,8,4,1,40,283,382]
d: [2,1420,1,6,3,2,21,1394,387]
d: [2,1421,1,4,2,3,21,1420,392]
d: [2,1517,1,12,6,0,44,326,276]
d: [2,1522,1,10,5,0,21,1517,279]
d: [2,1512,2,8,4,1,40,1449,324]
d: [2,2594,2,6,3,2,21,1512,387]
d: [2,2595,2,4,2,3,21,2594,392]
d: [2,2596,2,2,1,4,21,2595,1187]
f: 2596,2560,2531,1960,1919,1792,1522,1455,1436,1421,1393,1405,1094,1093,1042,1001,1000,966,687,681,352,353,354,267,264,254,252,251,248,247,240,241,242,243,244
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__l206_xreal_1
% SZS output start Proof for xreal_1__l206_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(k7_xcmplx_0(X0,X1),1)))) [input l206_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(k7_xcmplx_0(X0,X1),1)))) [negated conjecture 1]
7. ! [X0] : (m1_subset_1(X0,k1_numbers) => v1_xreal_0(X0)) [input cc1_xreal_0]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
13. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
14. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
19. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
40. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
41. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v2_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc31_xreal_0]
44. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc34_xreal_0]
48. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
50. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ! [X2] : (v1_xreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(k6_numbers,X0)) => r1_xxreal_0(k7_xcmplx_0(X1,X0),k7_xcmplx_0(X2,X0)))))) [input l66_xreal_1]
61. ? [X0] : (v1_xreal_0(X0) & v3_xxreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0)) [input rc3_xreal_0]
68. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
73. r1_xxreal_0(0,1) [input rqLessOrEqual__r1_xxreal_0__r0_r1]
77. v1_xboole_0(0) [input spc0_boole]
78. m1_subset_1(0,k1_numbers) & m1_subset_1(0,k5_numbers) & m2_subset_1(0,k1_numbers,k5_numbers) [input spc0_numerals]
80. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
81. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
84. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
90. ! [X0] : (v1_xcmplx_0(X0) => k6_numbers = k7_xcmplx_0(k6_numbers,X0)) [input t5_arithm]
93. ! [X0] : (v1_xcmplx_0(X0) => (k6_numbers != X0 => 1 = k7_xcmplx_0(X0,X0))) [input t60_xcmplx_1]
95. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
100. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
112. ? [X0] : (? [X1] : ((~r1_xxreal_0(k7_xcmplx_0(X0,X1),1) & (r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,X1))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
113. ? [X0] : (? [X1] : (~r1_xxreal_0(k7_xcmplx_0(X0,X1),1) & r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 112]
117. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
118. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
121. ! [X0] : (k6_numbers = k7_xcmplx_0(k6_numbers,X0) | ~v1_xcmplx_0(X0)) [ennf transformation 90]
123. ! [X0] : ((1 = k7_xcmplx_0(X0,X0) | k6_numbers = X0) | ~v1_xcmplx_0(X0)) [ennf transformation 93]
124. ! [X0] : (1 = k7_xcmplx_0(X0,X0) | k6_numbers = X0 | ~v1_xcmplx_0(X0)) [flattening 123]
125. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 13]
131. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 84]
132. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 131]
135. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 100]
136. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 135]
143. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(k7_xcmplx_0(X1,X0),k7_xcmplx_0(X2,X0)) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(k6_numbers,X0))) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 50]
144. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(k7_xcmplx_0(X1,X0),k7_xcmplx_0(X2,X0)) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(k6_numbers,X0) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 143]
148. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 95]
149. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
151. ! [X0] : (v1_xreal_0(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 7]
160. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 19]
161. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 160]
162. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 14]
163. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 162]
178. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 41]
179. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 178]
182. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 44]
183. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 182]
194. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 48]
195. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 194]
200. ? [X0] : (? [X1] : (~r1_xxreal_0(k7_xcmplx_0(X0,X1),1) & r1_xxreal_0(k6_numbers,X0) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~r1_xxreal_0(k7_xcmplx_0(sK0,X1),1) & r1_xxreal_0(k6_numbers,sK0) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
201. ? [X1] : (~r1_xxreal_0(k7_xcmplx_0(sK0,X1),1) & r1_xxreal_0(k6_numbers,sK0) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) => (~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),1) & r1_xxreal_0(k6_numbers,sK0) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) [choice axiom]
202. (~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),1) & r1_xxreal_0(k6_numbers,sK0) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 113,201,200]
228. ? [X0] : (v1_xreal_0(X0) & v3_xxreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0)) => (v1_xreal_0(sK14) & v3_xxreal_0(sK14) & v1_xxreal_0(sK14) & v1_xcmplx_0(sK14)) [choice axiom]
229. v1_xreal_0(sK14) & v3_xxreal_0(sK14) & v1_xxreal_0(sK14) & v1_xcmplx_0(sK14) [skolemisation 61,228]
240. v1_xreal_0(sK0) [cnf transformation 202] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
241. v1_xreal_0(sK1) [cnf transformation 202] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
242. r1_xxreal_0(sK0,sK1) [cnf transformation 202] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
243. r1_xxreal_0(k6_numbers,sK0) [cnf transformation 202] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
244. ~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),1) [cnf transformation 202] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
247. v1_xboole_0(k1_xboole_0) [cnf transformation 40] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
248. v1_xboole_0(0) [cnf transformation 77] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
251. k6_numbers = k1_xboole_0 [cnf transformation 68] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
252. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 81] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
254. r1_xxreal_0(0,1) [cnf transformation 73] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
264. m1_subset_1(1,k1_numbers) [cnf transformation 80] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
267. m1_subset_1(0,k1_numbers) [cnf transformation 78] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
271. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 117] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
272. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 118] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
276. k6_numbers = k7_xcmplx_0(k6_numbers,X0) | ~v1_xcmplx_0(X0) [cnf transformation 121] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
278. 1 = k7_xcmplx_0(X0,X0) | k6_numbers = X0 | ~v1_xcmplx_0(X0) [cnf transformation 124] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
279. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 125] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
283. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 132] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
285. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 136] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
289. r1_xxreal_0(k7_xcmplx_0(X1,X0),k7_xcmplx_0(X2,X0)) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(k6_numbers,X0) | ~v1_xreal_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 144] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
293. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 148] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
294. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 149] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
296. ~m1_subset_1(X0,k1_numbers) | v1_xreal_0(X0) [cnf transformation 151] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
306. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 161] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
307. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 163] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
324. ~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 179] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
326. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 183] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
333. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 195] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
352. v1_xxreal_0(sK14) [cnf transformation 229] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
353. v3_xxreal_0(sK14) [cnf transformation 229] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
354. v1_xreal_0(sK14) [cnf transformation 229] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
370. 1 <=> r1_xxreal_0(k7_xcmplx_0(sK0,sK1),1) [avatar definition]
372. ~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),1) <- (~1) [avatar component clause 370] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
373. ~1 [avatar split clause 244,370]
375. 2 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
377. r1_xxreal_0(k6_numbers,sK0) <- (2) [avatar component clause 375] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
378. 2 [avatar split clause 243,375]
380. 3 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
382. r1_xxreal_0(sK0,sK1) <- (3) [avatar component clause 380] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
383. 3 [avatar split clause 242,380]
385. 4 <=> v1_xreal_0(sK1) [avatar definition]
387. v1_xreal_0(sK1) <- (4) [avatar component clause 385] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
388. 4 [avatar split clause 241,385]
390. 5 <=> v1_xreal_0(sK0) [avatar definition]
392. v1_xreal_0(sK0) <- (5) [avatar component clause 390] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
393. 5 [avatar split clause 240,390]
405. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
407. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 405] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
408. 8 [avatar split clause 247,405]
410. 9 <=> v1_xboole_0(0) [avatar definition]
412. v1_xboole_0(0) <- (9) [avatar component clause 410] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
413. 9 [avatar split clause 248,410]
425. 12 <=> k6_numbers = k1_xboole_0 [avatar definition]
427. k6_numbers = k1_xboole_0 <- (12) [avatar component clause 425] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
428. 12 [avatar split clause 251,425]
430. 13 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
432. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (13) [avatar component clause 430] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
433. 13 [avatar split clause 252,430]
440. 15 <=> r1_xxreal_0(0,1) [avatar definition]
442. r1_xxreal_0(0,1) <- (15) [avatar component clause 440] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
443. 15 [avatar split clause 254,440]
475. 22 <=> m1_subset_1(1,k1_numbers) [avatar definition]
477. m1_subset_1(1,k1_numbers) <- (22) [avatar component clause 475] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
478. 22 [avatar split clause 264,475]
495. 26 <=> m1_subset_1(0,k1_numbers) [avatar definition]
497. m1_subset_1(0,k1_numbers) <- (26) [avatar component clause 495] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
498. 26 [avatar split clause 267,495]
592. 46 <=> v1_xreal_0(sK14) [avatar definition]
594. v1_xreal_0(sK14) <- (46) [avatar component clause 592] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
595. 46 [avatar split clause 354,592]
597. 47 <=> v3_xxreal_0(sK14) [avatar definition]
599. v3_xxreal_0(sK14) <- (47) [avatar component clause 597] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
600. 47 [avatar split clause 353,597]
602. 48 <=> v1_xxreal_0(sK14) [avatar definition]
604. v1_xxreal_0(sK14) <- (48) [avatar component clause 602] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
605. 48 [avatar split clause 352,602]
681. v1_xboole_0(k6_numbers) <- (8, 12) [backward demodulation 407,427] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
683. 64 <=> v1_xboole_0(k6_numbers) [avatar definition]
685. v1_xboole_0(k6_numbers) <- (64) [avatar component clause 683] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
686. 64 | ~8 | ~12 [avatar split clause 681,425,405,683]
687. m1_subset_1(k6_numbers,k4_ordinal1) <- (12, 13) [forward demodulation 432,427] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
689. 65 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
691. m1_subset_1(k6_numbers,k4_ordinal1) <- (65) [avatar component clause 689] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
692. 65 | ~12 | ~13 [avatar split clause 687,430,425,689]
881. 95 <=> v7_ordinal1(k6_numbers) [avatar definition]
883. v7_ordinal1(k6_numbers) <- (95) [avatar component clause 881] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
965. k6_numbers = X0 | ~v1_xboole_0(X0) <- (12) [forward demodulation 293,427] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
966. v7_ordinal1(k6_numbers) <- (65) [resolution 294,691] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
970. 95 | ~65 [avatar split clause 966,689,881]
1000. v1_xreal_0(1) <- (22) [resolution 296,477] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1001. v1_xreal_0(0) <- (26) [resolution 296,497] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1004. 114 <=> v1_xreal_0(1) [avatar definition]
1007. 114 | ~22 [avatar split clause 1000,475,1004]
1009. 115 <=> v1_xreal_0(0) [avatar definition]
1012. 115 | ~26 [avatar split clause 1001,495,1009]
1040. ~v2_xxreal_0(sK14) | ~v1_xxreal_0(sK14) <- (47) [resolution 306,599] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:3,thDist:-3}
1042. ~v2_xxreal_0(sK14) <- (47, 48) [subsumption resolution 1040,604] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:3,thDist:-3}
1044. 119 <=> v2_xxreal_0(sK14) [avatar definition]
1046. ~v2_xxreal_0(sK14) <- (~119) [avatar component clause 1044] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
1047. ~119 | ~47 | ~48 [avatar split clause 1042,602,597,1044]
1093. v1_xxreal_0(k6_numbers) <- (95) [resolution 883,272] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1094. v1_xreal_0(k6_numbers) <- (95) [resolution 883,271] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1099. 125 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1101. v1_xxreal_0(k6_numbers) <- (125) [avatar component clause 1099] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
1102. 125 | ~95 [avatar split clause 1093,881,1099]
1104. 126 <=> v1_xreal_0(k6_numbers) [avatar definition]
1106. v1_xreal_0(k6_numbers) <- (126) [avatar component clause 1104] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
1107. 126 | ~95 [avatar split clause 1094,881,1104]
1186. 136 <=> v3_xxreal_0(sK0) [avatar definition]
1187. ~v3_xxreal_0(sK0) <- (~136) [avatar component clause 1186] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1195. 138 <=> v2_xxreal_0(sK1) [avatar definition]
1199. 139 <=> v3_xxreal_0(sK1) [avatar definition]
1200. ~v3_xxreal_0(sK1) <- (~139) [avatar component clause 1199] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1390. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (2) [resolution 283,377] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1393. ~v3_xxreal_0(1) | v3_xxreal_0(0) | ~v1_xreal_0(1) | ~v1_xreal_0(0) <- (15) [resolution 283,442] {a:1,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
1394. ~v3_xxreal_0(sK1) | v3_xxreal_0(sK0) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (3) [resolution 283,382] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1405. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 5) [subsumption resolution 1390,392] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1407. 170 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1408. ~v3_xxreal_0(k6_numbers) <- (~170) [avatar component clause 1407] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1410. ~126 | 170 | ~136 | ~2 | ~5 [avatar split clause 1405,390,375,1186,1407,1104]
1412. 171 <=> v3_xxreal_0(0) [avatar definition]
1414. v3_xxreal_0(0) <- (171) [avatar component clause 1412] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1416. 172 <=> v3_xxreal_0(1) [avatar definition]
1419. ~115 | ~114 | 171 | ~172 | ~15 [avatar split clause 1393,440,1416,1412,1004,1009]
1420. ~v3_xxreal_0(sK1) | v3_xxreal_0(sK0) | ~v1_xreal_0(sK0) <- (3, 4) [subsumption resolution 1394,387] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1421. ~v3_xxreal_0(sK1) | v3_xxreal_0(sK0) <- (3, 4, 5) [subsumption resolution 1420,392] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1422. 136 | ~139 | ~3 | ~4 | ~5 [avatar split clause 1421,390,385,380,1199,1186]
1436. v3_xxreal_0(1) | v2_xxreal_0(k7_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(1) | ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~1) [resolution 285,372] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1443. 173 <=> v1_xreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1445. ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~173) [avatar component clause 1443] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1447. 174 <=> v2_xxreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1449. v2_xxreal_0(k7_xcmplx_0(sK0,sK1)) <- (174) [avatar component clause 1447] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1450. ~173 | ~114 | 174 | 172 | 1 [avatar split clause 1436,370,1416,1447,1004,1443]
1452. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~173) [resolution 1445,333] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1454. ~v1_xreal_0(sK0) <- (4, ~173) [subsumption resolution 1452,387] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1455. $false <- (4, 5, ~173) [subsumption resolution 1454,392] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1456. ~4 | ~5 | 173 [avatar contradiction clause 1455]
1484. 177 <=> ! [X0] : (~v1_xreal_0(X0) | v2_xxreal_0(X0)) [avatar definition]
1485. ~v1_xreal_0(X0) | v2_xxreal_0(X0) <- (177) [avatar component clause 1484] {a:1,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1487. 178 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1488. v2_xxreal_0(k6_numbers) <- (178) [avatar component clause 1487] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1489. ~v2_xxreal_0(k6_numbers) <- (~178) [avatar component clause 1487] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1512. ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (174) [resolution 1449,324] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1517. ~v3_xxreal_0(k6_numbers) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) | ~v1_xreal_0(k6_numbers) | v2_xxreal_0(k6_numbers) | ~v1_xcmplx_0(X0) [superposition 326,276] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
1522. ~v3_xxreal_0(k6_numbers) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) | ~v1_xreal_0(k6_numbers) | v2_xxreal_0(k6_numbers) [subsumption resolution 1517,279] {a:1,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
1523. 178 | ~126 | 177 | ~170 [avatar split clause 1522,1407,1484,1104,1487]
1537. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) <- (12, ~170) [superposition 1408,965] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1625. r1_xxreal_0(k7_xcmplx_0(X6,X5),1) | ~r1_xxreal_0(X6,X5) | ~r1_xxreal_0(k6_numbers,X5) | ~v1_xreal_0(X5) | ~v1_xreal_0(X6) | ~v1_xreal_0(X5) | k6_numbers = X5 | ~v1_xcmplx_0(X5) [superposition 289,278] {a:1,w:22,wCS:8800,sine:1,thAx:0,allAx:2,thDist:-2}
1626. r1_xxreal_0(k7_xcmplx_0(X6,X5),1) | ~r1_xxreal_0(X6,X5) | ~r1_xxreal_0(k6_numbers,X5) | ~v1_xreal_0(X5) | ~v1_xreal_0(X6) | k6_numbers = X5 | ~v1_xcmplx_0(X5) [duplicate literal removal 1625] {a:1,w:20,wCS:8000,sine:1,thAx:0,allAx:2,thDist:-2}
1671. r1_xxreal_0(k7_xcmplx_0(X6,X5),1) | ~r1_xxreal_0(X6,X5) | ~r1_xxreal_0(k6_numbers,X5) | ~v1_xreal_0(X5) | ~v1_xreal_0(X6) | k6_numbers = X5 [subsumption resolution 1626,279] {a:1,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1789. ~v1_xboole_0(0) <- (12, ~170, 171) [resolution 1414,1537] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1792. $false <- (9, 12, ~170, 171) [subsumption resolution 1789,412] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1793. ~9 | ~12 | 170 | ~171 [avatar contradiction clause 1792]
1915. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (178) [resolution 1488,307] {a:2,w:4,wCS:1600,sine:2,thAx:0,allAx:3,thDist:-3}
1918. ~v1_xxreal_0(k6_numbers) <- (64, 178) [subsumption resolution 1915,685] {a:2,w:2,wCS:800,sine:2,thAx:0,allAx:3,thDist:-3}
1919. $false <- (64, 125, 178) [subsumption resolution 1918,1101] {a:2,w:0,sine:2,thAx:0,allAx:3,thDist:-3}
1920. ~64 | ~125 | ~178 [avatar contradiction clause 1919]
1953. v2_xxreal_0(sK14) <- (46, 177) [resolution 1485,594] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:4,thDist:-4}
1960. $false <- (46, ~119, 177) [subsumption resolution 1953,1046] {a:2,w:0,sine:255,thAx:0,allAx:4,thDist:-4}
1961. ~46 | 119 | ~177 [avatar contradiction clause 1960]
2516. ~r1_xxreal_0(sK0,sK1) | ~r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) | k6_numbers = sK1 <- (~1) [resolution 1671,372] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2529. ~r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) | k6_numbers = sK1 <- (~1, 3) [subsumption resolution 2516,382] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2530. ~r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(sK0) | k6_numbers = sK1 <- (~1, 3, 4) [subsumption resolution 2529,387] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2531. ~r1_xxreal_0(k6_numbers,sK1) | k6_numbers = sK1 <- (~1, 3, 4, 5) [subsumption resolution 2530,392] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2533. 236 <=> k6_numbers = sK1 [avatar definition]
2537. 237 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
2539. ~r1_xxreal_0(k6_numbers,sK1) <- (~237) [avatar component clause 2537] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2540. 236 | ~237 | 1 | ~3 | ~4 | ~5 [avatar split clause 2531,390,385,380,370,2537,2533]
2546. v3_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (~237) [resolution 2539,285] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2557. v2_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (~139, ~237) [subsumption resolution 2546,1200] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2558. ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (~139, ~178, ~237) [subsumption resolution 2557,1489] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2559. ~v1_xreal_0(k6_numbers) <- (4, ~139, ~178, ~237) [subsumption resolution 2558,387] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2560. $false <- (4, 126, ~139, ~178, ~237) [subsumption resolution 2559,1106] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2561. ~4 | ~126 | 139 | 178 | 237 [avatar contradiction clause 2560]
2565. k6_numbers != sK1 | v2_xxreal_0(k6_numbers) | ~v2_xxreal_0(sK1) [theory tautology sat conflict] {a:0,w:7,wCS:2800,sine:255,thAx:0,allAx:1,thDist:-1}
2594. v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (4, 174) [subsumption resolution 1512,387] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2595. v2_xxreal_0(sK1) | v3_xxreal_0(sK0) <- (4, 5, 174) [subsumption resolution 2594,392] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2596. v2_xxreal_0(sK1) <- (4, 5, ~136, 174) [subsumption resolution 2595,1187] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2597. 138 | ~4 | ~5 | 136 | ~174 [avatar split clause 2596,1447,1186,390,385,1195]
2602. $false [avatar sat refutation 373,378,383,388,393,408,413,428,433,443,478,498,595,600,605,686,692,970,1007,1012,1047,1102,1107,1410,1419,1422,1450,1456,1523,1793,1920,1961,2540,2561,2565,2597] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1546,thDist:-1546}
% SZS output end Proof for xreal_1__l206_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 99
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 129
% Generated clauses: 1400
% Active clauses: 291
% Passive clauses: 432
% Final active clauses: 260
% Final passive clauses: 114
% Main loop iterations started: 312
% 
% >>> Simplifying Inferences
% Duplicate literals: 146
% Fw subsumption resolutions: 195
% Fw demodulations: 8
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 96
% Forward subsumptions: 277
% Backward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 523
% Unit resulting resolution: 12
% Factoring: 1
% Forward superposition: 125
% Backward superposition: 45
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 47
% Split components: 119
% Sat splitting refutations: 24
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 288
% SAT solver unit clauses: 64
% SAT solver binary clauses: 126
% 
% Memory used [KB]: 6652
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% term sharing: 0.001 s
% non unit clause index maintenance: 0.001 s ( own 0.0000 s ) 
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.002 s
% unit resulting resolution: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.010 s ( own 0.003 s ) 

