i: [1,46,0,2,1,1,0,0] t6_necklace
a: [3,76,0,2,1,46]
s: 76
i: [1,45,0,2,1,1,0,0] t6_necklace
a: [3,81,0,2,1,45]
s: 81
i: [1,44,0,2,1,1,0,0] t6_necklace
a: [3,86,0,2,1,44]
s: 86
i: [1,43,0,2,1,1,0,0] t6_necklace
a: [3,91,0,2,1,43]
s: 91
i: [1,47,0,5,1,1,0,0] t6_necklace
a: [3,71,0,5,1,47]
s: 71
i: [1,62,0,2,1,0,0,1] rc1_funct_1
a: [3,107,0,2,1,62]
s: 107
i: [1,61,0,2,1,0,0,1] rc1_funct_1
a: [3,112,0,2,1,61]
s: 112
i: [1,48,0,10,1,1,0,0] t6_necklace
a: [3,66,0,10,1,48]
s: 66
i: [1,51,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 51
i: [1,53,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 53
i: [1,52,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 52
d: [2,114,1,5,1,1,40,53,71]
a: [3,118,1,5,1,114]
s: 118
i: [1,54,0,11,5,0,0,2] reflexivity_r1_partfun1
a: [3,98,0,7,3,54]
s: 98
i: [1,60,0,8,3,0,0,1] fc3_relat_1
s: 60
i: [1,55,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,102,0,9,3,55]
s: 102
i: [1,57,0,12,5,0,0,1] dt_k1_funct_4
s: 57
i: [1,58,0,12,5,0,0,1] dt_k1_funct_4
s: 58
i: [1,59,0,14,6,0,0,2] symmetry_r1_partfun1
s: 59
i: [1,50,0,16,6,0,0,1] t56_partfun1
s: 50
i: [1,56,0,17,5,0,0,1] rd3_funct_4
s: 56
i: [1,49,0,21,6,0,0,1] t35_frechet
s: 49
d: [2,131,1,11,5,1,40,50,71]
d: [2,133,1,9,4,2,21,131,76]
d: [2,134,1,7,3,3,21,133,81]
d: [2,135,1,5,2,4,21,134,86]
d: [2,136,1,3,1,5,21,135,91]
d: [2,162,1,20,6,1,44,66,49]
d: [2,168,1,11,5,1,19,162]
d: [2,171,1,9,4,2,21,168,76]
d: [2,172,1,7,3,3,21,171,81]
d: [2,173,1,5,2,4,21,172,86]
d: [2,174,1,3,1,5,21,173,91]
f: 174,136,43,44,45,46,47,48
% Refutation found. Thanks to Tanya!
% SZS status Theorem for necklace__t6_necklace
% SZS output start Proof for necklace__t6_necklace
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k10_xtuple_0(k1_funct_4(X0,X1)) = k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1))))) [input t6_necklace]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k10_xtuple_0(k1_funct_4(X0,X1)) = k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1))))) [negated conjecture 1]
16. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) => k10_xtuple_0(k1_funct_4(X0,X1)) = k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1))))) [input t35_frechet]
17. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => r1_partfun1(X0,X1)))) [input t56_partfun1]
19. ? [X0] : (? [X1] : ((k10_xtuple_0(k1_funct_4(X0,X1)) != k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
20. ? [X0] : (? [X1] : (k10_xtuple_0(k1_funct_4(X0,X1)) != k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 19]
21. ! [X0] : (! [X1] : ((k10_xtuple_0(k1_funct_4(X0,X1)) = k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_partfun1(X0,X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 16]
22. ! [X0] : (! [X1] : (k10_xtuple_0(k1_funct_4(X0,X1)) = k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 21]
23. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 17]
24. ! [X0] : (! [X1] : (r1_partfun1(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 23]
38. ? [X0] : (? [X1] : (k10_xtuple_0(k1_funct_4(X0,X1)) != k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (k10_xtuple_0(k1_funct_4(sK0,X1)) != k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(X1)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
39. ? [X1] : (k10_xtuple_0(k1_funct_4(sK0,X1)) != k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(X1)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) => (k10_xtuple_0(k1_funct_4(sK0,sK1)) != k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
40. (k10_xtuple_0(k1_funct_4(sK0,sK1)) != k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 20,39,38]
43. v1_relat_1(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
44. v1_funct_1(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
45. v1_relat_1(sK1) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
46. v1_funct_1(sK1) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
47. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) [cnf transformation 40] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. k10_xtuple_0(k1_funct_4(sK0,sK1)) != k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) [cnf transformation 40] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. k10_xtuple_0(k1_funct_4(X0,X1)) = k2_xboole_0(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. 1 <=> k10_xtuple_0(k1_funct_4(sK0,sK1)) = k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) [avatar definition]
66. k10_xtuple_0(k1_funct_4(sK0,sK1)) != k2_xboole_0(k10_xtuple_0(sK0),k10_xtuple_0(sK1)) <- (~1) [avatar component clause 64] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
67. ~1 [avatar split clause 48,64]
69. 2 <=> r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) [avatar definition]
71. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) <- (2) [avatar component clause 69] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. 2 [avatar split clause 47,69]
74. 3 <=> v1_funct_1(sK1) [avatar definition]
76. v1_funct_1(sK1) <- (3) [avatar component clause 74] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
77. 3 [avatar split clause 46,74]
79. 4 <=> v1_relat_1(sK1) [avatar definition]
81. v1_relat_1(sK1) <- (4) [avatar component clause 79] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. 4 [avatar split clause 45,79]
84. 5 <=> v1_funct_1(sK0) [avatar definition]
86. v1_funct_1(sK0) <- (5) [avatar component clause 84] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
87. 5 [avatar split clause 44,84]
89. 6 <=> v1_relat_1(sK0) [avatar definition]
91. v1_relat_1(sK0) <- (6) [avatar component clause 89] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. 6 [avatar split clause 43,89]
131. r1_partfun1(sK0,sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 50,71] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
133. r1_partfun1(sK0,sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3) [subsumption resolution 131,76] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
134. r1_partfun1(sK0,sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3, 4) [subsumption resolution 133,81] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
135. r1_partfun1(sK0,sK1) | ~v1_relat_1(sK0) <- (2, 3, 4, 5) [subsumption resolution 134,86] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
136. r1_partfun1(sK0,sK1) <- (2, 3, 4, 5, 6) [subsumption resolution 135,91] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
138. 13 <=> r1_partfun1(sK0,sK1) [avatar definition]
141. 13 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar split clause 136,89,84,79,74,69,138]
162. k10_xtuple_0(k1_funct_4(sK0,sK1)) != k10_xtuple_0(k1_funct_4(sK0,sK1)) | ~r1_partfun1(sK0,sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [superposition 66,49] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
168. ~r1_partfun1(sK0,sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [trivial inequality removal 162] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
171. ~r1_partfun1(sK0,sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 3) [subsumption resolution 168,76] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
172. ~r1_partfun1(sK0,sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 3, 4) [subsumption resolution 171,81] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
173. ~r1_partfun1(sK0,sK1) | ~v1_relat_1(sK0) <- (~1, 3, 4, 5) [subsumption resolution 172,86] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
174. ~r1_partfun1(sK0,sK1) <- (~1, 3, 4, 5, 6) [subsumption resolution 173,91] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
175. ~13 | 1 | ~3 | ~4 | ~5 | ~6 [avatar split clause 174,89,84,79,74,64,138]
177. $false [avatar sat refutation 67,72,77,82,87,92,141,175] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:73,thDist:-73}
% SZS output end Proof for necklace__t6_necklace
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 20
% Generated clauses: 80
% Active clauses: 21
% Passive clauses: 30
% Final active clauses: 21
% Final passive clauses: 9
% Main loop iterations started: 21
% 
% >>> Simplifying Inferences
% Duplicate literals: 27
% Trivial inequalities: 1
% Fw subsumption resolutions: 14
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 5
% Forward superposition: 9
% Backward superposition: 7
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 8
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

