i: [1,45,0,3,1,1,0,0] t17_zfmisc_1
a: [3,81,0,3,1,45]
s: 81
i: [1,47,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,86,0,2,1,47]
s: 86
i: [1,48,0,2,1,0,0,2] fc1_xboole_0
a: [3,91,0,2,1,48]
s: 91
i: [1,67,0,2,1,0,0,2] rc2_xboole_0
a: [3,101,0,2,1,67]
s: 101
i: [1,68,0,2,1,0,0,2] rc1_xboole_0
a: [3,106,0,2,1,68]
s: 106
i: [1,46,0,9,1,1,0,0] t17_zfmisc_1
a: [3,76,0,9,1,46]
s: 76
i: [1,50,0,3,1,0,0,1] fc2_xboole_0
s: 50
i: [1,49,0,3,1,0,0,255] d2_xboole_0
a: [3,96,0,3,1,49]
s: 96
r: 86
i: [1,64,0,4,1,0,0,1] fc3_xboole_0
s: 64
i: [1,70,0,4,1,0,0,1] d1_tarski
s: 70
i: [1,51,0,5,1,0,0,1] t4_boole
s: 51
i: [1,52,0,5,1,0,0,1] t3_boole
s: 52
i: [1,53,0,5,2,0,0,2] t6_boole
s: 53
i: [1,59,0,5,2,0,0,2] t7_boole
s: 59
i: [1,66,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 66
i: [1,58,0,7,3,0,0,1] t8_boole
s: 58
d: [2,114,1,4,1,0,40,66,70]
s: 114
i: [1,65,0,7,1,0,0,1] commutativity_k2_tarski
s: 65
i: [1,71,0,7,2,0,0,1] d1_tarski
s: 71
i: [1,60,0,11,2,0,0,1] l31_zfmisc_1
s: 60
i: [1,72,0,11,2,0,0,1] l31_zfmisc_1
s: 72
d: [2,109,1,7,3,0,44,53,53]
s: 109
i: [1,56,0,14,3,0,0,1] d1_tarski
s: 56
d: [2,110,1,7,2,0,44,52,53]
s: 110
i: [1,57,0,14,3,0,0,1] d1_tarski
s: 57
i: [1,61,0,14,3,0,0,1] l31_zfmisc_1
s: 61
d: [2,127,2,11,3,0,44,110,72]
d: [2,131,2,8,2,0,21,127,59]
a: [3,134,2,2,1,131]
s: 134
r: 101
r: 50
r: 64
r: 53
r: 59
r: 58
r: 110
d: [2,153,3,0,0,2,40,134,91]
e: 153
d: [2,152,3,0,0,2,40,134,106]
e: 152
d: [2,151,3,0,0,2,52,106,134]
e: 151
r: 134
s: 101
s: 50
s: 64
s: 53
s: 59
s: 58
a: [3,137,2,6,1,131]
s: 137
r: 72
d: [2,165,0,10,2,1,23,72,137]
s: 165
i: [1,62,0,14,3,0,0,1] l31_zfmisc_1
s: 62
d: [2,173,1,13,3,1,44,76,62]
d: [2,177,1,8,2,1,19,173]
d: [2,178,1,4,1,1,21,177,70]
a: [3,182,1,4,1,178]
s: 182
d: [2,184,2,0,0,2,52,81,182,71]
e: 184
d: [2,185,2,3,1,1,40,182,71]
d: [2,189,2,0,0,2,21,185,81]
e: 189
f: 189,178,45,46
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t17_zfmisc_1
% SZS output start Proof for zfmisc_1__t17_zfmisc_1
1. ! [X0] : ! [X1] : (X0 != X1 => k4_xboole_0(k2_tarski(X0,X1),k1_tarski(X1)) = k1_tarski(X0)) [input t17_zfmisc_1]
2. ~! [X0] : ! [X1] : (X0 != X1 => k4_xboole_0(k2_tarski(X0,X1),k1_tarski(X1)) = k1_tarski(X0)) [negated conjecture 1]
5. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
15. ! [X0] : ! [X1] : ! [X2] : (k1_tarski(X0) = k4_xboole_0(k2_tarski(X0,X1),X2) <=> ((X0 = X1 | r2_hidden(X1,X2)) & ~r2_hidden(X0,X2))) [input l31_zfmisc_1]
23. ~! [X0,X1] : (X0 != X1 => k4_xboole_0(k2_tarski(X0,X1),k1_tarski(X1)) = k1_tarski(X0)) [flattening 2]
24. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 5]
27. ! [X0,X1,X2] : (k1_tarski(X0) = k4_xboole_0(k2_tarski(X0,X1),X2) <=> ((X0 = X1 | r2_hidden(X1,X2)) & ~r2_hidden(X0,X2))) [flattening 15]
28. ? [X0,X1] : (k4_xboole_0(k2_tarski(X0,X1),k1_tarski(X1)) != k1_tarski(X0) & X0 != X1) [ennf transformation 23]
33. ? [X0,X1] : (k4_xboole_0(k2_tarski(X0,X1),k1_tarski(X1)) != k1_tarski(X0) & X0 != X1) => (k4_xboole_0(k2_tarski(sK0,sK1),k1_tarski(sK1)) != k1_tarski(sK0) & sK0 != sK1) [choice axiom]
34. k4_xboole_0(k2_tarski(sK0,sK1),k1_tarski(sK1)) != k1_tarski(sK0) & sK0 != sK1 [skolemisation 28,33]
35. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 24]
36. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 35]
37. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) [choice axiom]
38. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 36,37]
39. ! [X0,X1,X2] : ((k1_tarski(X0) = k4_xboole_0(k2_tarski(X0,X1),X2) | ((X0 != X1 & ~r2_hidden(X1,X2)) | r2_hidden(X0,X2))) & (((X0 = X1 | r2_hidden(X1,X2)) & ~r2_hidden(X0,X2)) | k1_tarski(X0) != k4_xboole_0(k2_tarski(X0,X1),X2))) [nnf transformation 27]
40. ! [X0,X1,X2] : ((k1_tarski(X0) = k4_xboole_0(k2_tarski(X0,X1),X2) | (X0 != X1 & ~r2_hidden(X1,X2)) | r2_hidden(X0,X2)) & (((X0 = X1 | r2_hidden(X1,X2)) & ~r2_hidden(X0,X2)) | k1_tarski(X0) != k4_xboole_0(k2_tarski(X0,X1),X2))) [flattening 39]
45. sK0 != sK1 [cnf transformation 34] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. k4_xboole_0(k2_tarski(sK0,sK1),k1_tarski(sK1)) != k1_tarski(sK0) [cnf transformation 34] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 38] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
55. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 38] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
62. k1_tarski(X0) = k4_xboole_0(k2_tarski(X0,X1),X2) | ~r2_hidden(X1,X2) | r2_hidden(X0,X2) [cnf transformation 40] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 55] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:2,thDist:-2}
70. r2_hidden(X3,k1_tarski(X3)) [equality resolution 69] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
71. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 54] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
74. 1 <=> k4_xboole_0(k2_tarski(sK0,sK1),k1_tarski(sK1)) = k1_tarski(sK0) [avatar definition]
76. k4_xboole_0(k2_tarski(sK0,sK1),k1_tarski(sK1)) != k1_tarski(sK0) <- (~1) [avatar component clause 74] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
77. ~1 [avatar split clause 46,74]
79. 2 <=> sK0 = sK1 [avatar definition]
81. sK0 != sK1 <- (~2) [avatar component clause 79] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
82. ~2 [avatar split clause 45,79]
173. k1_tarski(sK0) != k1_tarski(sK0) | ~r2_hidden(sK1,k1_tarski(sK1)) | r2_hidden(sK0,k1_tarski(sK1)) <- (~1) [superposition 76,62] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
177. ~r2_hidden(sK1,k1_tarski(sK1)) | r2_hidden(sK0,k1_tarski(sK1)) <- (~1) [trivial inequality removal 173] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
178. r2_hidden(sK0,k1_tarski(sK1)) <- (~1) [subsumption resolution 177,70] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
180. 11 <=> r2_hidden(sK0,k1_tarski(sK1)) [avatar definition]
182. r2_hidden(sK0,k1_tarski(sK1)) <- (11) [avatar component clause 180] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
183. 11 | 1 [avatar split clause 178,74,180]
185. sK0 = sK1 <- (11) [resolution 182,71] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
189. $false <- (~2, 11) [subsumption resolution 185,81] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
190. 2 | ~11 [avatar contradiction clause 189]
191. $false [avatar sat refutation 77,82,183,190] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:64,thDist:-64}
% SZS output end Proof for zfmisc_1__t17_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 101
% Active clauses: 36
% Passive clauses: 60
% Final active clauses: 26
% Final passive clauses: 23
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 5
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Equational tautologies: 1
% Forward subsumptions: 12
% Backward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 2
% Forward superposition: 15
% Backward superposition: 19
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 7
% SAT solver binary clauses: 9
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

