i: [1,19,0,2,1,1,0,0] t8_finset_1
a: [3,39,0,2,1,19]
s: 39
i: [1,18,0,2,1,1,0,0] t8_finset_1
a: [3,44,0,2,1,18]
s: 44
i: [1,21,0,3,1,1,0,0] t8_finset_1
a: [3,29,0,3,1,21]
s: 29
i: [1,20,0,3,1,1,0,0] t8_finset_1
a: [3,34,0,3,1,20]
s: 34
i: [1,25,0,2,1,0,0,1] rc1_funct_1
a: [3,49,0,2,1,25]
s: 49
i: [1,24,0,2,1,0,0,1] rc1_funct_1
a: [3,54,0,2,1,24]
s: 54
i: [1,22,0,9,2,0,0,255] t113_relat_1
s: 22
i: [1,23,0,10,4,0,0,1] fc13_finset_1
s: 23
d: [2,56,1,12,5,0,44,23,22]
d: [2,57,1,10,4,0,18,56]
s: 57
d: [2,58,2,0,0,4,52,44,39,29,34,57]
e: 58
d: [2,59,2,7,3,1,40,57,34]
d: [2,61,2,4,2,2,21,59,29]
d: [2,62,2,2,1,3,21,61,39]
d: [2,63,2,0,0,4,21,62,44]
e: 63
f: 63,18,19,20,21
% Refutation found. Thanks to Tanya!
% SZS status Theorem for finset_1__t8_finset_1
% SZS output start Proof for finset_1__t8_finset_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_finset_1(k9_xtuple_0(X0)) => v1_finset_1(k10_xtuple_0(X0)))) [input t8_finset_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_finset_1(k9_xtuple_0(X0)) => v1_finset_1(k10_xtuple_0(X0)))) [negated conjecture 1]
6. ! [X0,X1] : ((v1_finset_1(X1) & v1_funct_1(X0) & v1_relat_1(X0)) => v1_finset_1(k7_relat_1(X0,X1))) [input fc13_finset_1]
8. ! [X0] : (v1_relat_1(X0) => k10_xtuple_0(X0) = k7_relat_1(X0,k9_xtuple_0(X0))) [input t113_relat_1]
9. ? [X0] : ((~v1_finset_1(k10_xtuple_0(X0)) & v1_finset_1(k9_xtuple_0(X0))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
10. ? [X0] : (~v1_finset_1(k10_xtuple_0(X0)) & v1_finset_1(k9_xtuple_0(X0)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 9]
11. ! [X0] : (k10_xtuple_0(X0) = k7_relat_1(X0,k9_xtuple_0(X0)) | ~v1_relat_1(X0)) [ennf transformation 8]
12. ! [X0,X1] : (v1_finset_1(k7_relat_1(X0,X1)) | (~v1_finset_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 6]
13. ! [X0,X1] : (v1_finset_1(k7_relat_1(X0,X1)) | ~v1_finset_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 12]
14. ? [X0] : (~v1_finset_1(k10_xtuple_0(X0)) & v1_finset_1(k9_xtuple_0(X0)) & v1_funct_1(X0) & v1_relat_1(X0)) => (~v1_finset_1(k10_xtuple_0(sK0)) & v1_finset_1(k9_xtuple_0(sK0)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
15. ~v1_finset_1(k10_xtuple_0(sK0)) & v1_finset_1(k9_xtuple_0(sK0)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 10,14]
18. v1_relat_1(sK0) [cnf transformation 15] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
19. v1_funct_1(sK0) [cnf transformation 15] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
20. v1_finset_1(k9_xtuple_0(sK0)) [cnf transformation 15] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
21. ~v1_finset_1(k10_xtuple_0(sK0)) [cnf transformation 15] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
22. k10_xtuple_0(X0) = k7_relat_1(X0,k9_xtuple_0(X0)) | ~v1_relat_1(X0) [cnf transformation 11] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
23. v1_finset_1(k7_relat_1(X0,X1)) | ~v1_finset_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 13] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
27. 1 <=> v1_finset_1(k10_xtuple_0(sK0)) [avatar definition]
29. ~v1_finset_1(k10_xtuple_0(sK0)) <- (~1) [avatar component clause 27] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
30. ~1 [avatar split clause 21,27]
32. 2 <=> v1_finset_1(k9_xtuple_0(sK0)) [avatar definition]
34. v1_finset_1(k9_xtuple_0(sK0)) <- (2) [avatar component clause 32] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
35. 2 [avatar split clause 20,32]
37. 3 <=> v1_funct_1(sK0) [avatar definition]
39. v1_funct_1(sK0) <- (3) [avatar component clause 37] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
40. 3 [avatar split clause 19,37]
42. 4 <=> v1_relat_1(sK0) [avatar definition]
44. v1_relat_1(sK0) <- (4) [avatar component clause 42] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
45. 4 [avatar split clause 18,42]
56. v1_finset_1(k10_xtuple_0(X0)) | ~v1_finset_1(k9_xtuple_0(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v1_relat_1(X0) [superposition 23,22] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
57. ~v1_finset_1(k9_xtuple_0(X0)) | v1_finset_1(k10_xtuple_0(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [duplicate literal removal 56] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
59. v1_finset_1(k10_xtuple_0(sK0)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 57,34] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
61. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 59,29] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
62. ~v1_relat_1(sK0) <- (~1, 2, 3) [subsumption resolution 61,39] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
63. $false <- (~1, 2, 3, 4) [subsumption resolution 62,44] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
64. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 63]
65. $false [avatar sat refutation 30,35,40,45,64] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:32,thDist:-32}
% SZS output end Proof for finset_1__t8_finset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 7
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 8
% Generated clauses: 21
% Active clauses: 9
% Passive clauses: 9
% Final active clauses: 9
% Main loop iterations started: 9
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 3
% 
% >>> Generating Inferences
% Binary resolution: 1
% Unit resulting resolution: 1
% Forward superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 6
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

