i: [1,54,0,6,1,1,0,0] t7_zfmisc_1
a: [3,92,0,6,1,54]
s: 92
i: [1,55,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,97,0,2,1,55]
s: 97
i: [1,56,0,2,1,0,0,255] fc1_xboole_0
a: [3,102,0,2,1,56]
s: 102
i: [1,79,0,2,1,0,0,2] rc2_xboole_0
a: [3,112,0,2,1,79]
s: 112
i: [1,80,0,2,1,0,0,2] rc1_xboole_0
a: [3,117,0,2,1,80]
s: 117
i: [1,58,0,3,1,0,0,1] fc2_xboole_0
s: 58
i: [1,75,0,3,1,0,0,1] reflexivity_r1_tarski
s: 75
i: [1,57,0,3,1,0,0,255] d2_xboole_0
a: [3,107,0,3,1,57]
s: 107
r: 97
i: [1,76,0,4,1,0,0,1] fc3_xboole_0
s: 76
i: [1,82,0,4,1,0,0,1] d1_tarski
s: 82
i: [1,59,0,5,2,0,0,255] t6_boole
s: 59
i: [1,68,0,5,2,0,0,2] t7_boole
s: 68
i: [1,85,0,5,1,0,0,1] d2_tarski
s: 85
i: [1,87,0,5,1,0,0,1] d2_tarski
s: 87
i: [1,78,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 78
i: [1,67,0,7,3,0,0,2] t8_boole
s: 67
d: [2,125,1,4,1,0,40,78,82]
s: 125
i: [1,77,0,7,1,0,0,1] commutativity_k2_tarski
s: 77
d: [2,126,1,5,1,0,40,78,85]
s: 126
i: [1,83,0,7,2,0,0,1] d1_tarski
s: 83
i: [1,61,0,8,2,0,0,1] d3_tarski
s: 61
d: [2,127,1,5,1,0,40,78,87]
s: 127
i: [1,62,0,8,2,0,0,1] d3_tarski
s: 62
d: [2,138,1,5,2,0,40,61,68]
s: 138
i: [1,60,0,9,3,0,0,1] d3_tarski
s: 60
i: [1,88,0,11,3,0,0,1] d2_tarski
s: 88
d: [2,134,2,5,1,0,44,126,77]
s: 134
i: [1,65,0,14,3,0,0,1] d1_tarski
s: 65
d: [2,135,2,5,1,0,44,126,77]
s: 135
i: [1,66,0,14,3,0,0,1] d1_tarski
s: 66
i: [1,73,0,17,3,0,0,1] d2_tarski
s: 73
d: [2,120,1,7,3,0,44,59,59]
s: 120
i: [1,74,0,17,3,0,0,1] d2_tarski
s: 74
d: [2,145,1,7,2,0,40,60,82]
s: 145
i: [1,72,0,23,4,0,0,1] d2_tarski
s: 72
d: [2,137,1,8,2,0,40,61,78]
s: 137
d: [2,146,1,8,2,0,40,60,85]
s: 146
d: [2,147,1,8,2,0,40,60,87]
s: 147
d: [2,168,2,8,2,0,44,146,77]
s: 168
d: [2,139,1,10,2,0,40,61,83]
s: 139
d: [2,148,1,11,3,0,40,60,61]
s: 148
d: [2,175,2,11,3,0,44,62,139]
d: [2,178,2,7,2,0,18,175]
s: 178
d: [2,190,3,0,0,1,52,87,92,178]
e: 190
d: [2,191,3,5,1,1,40,178,92]
d: [2,194,3,0,0,1,21,191,87]
e: 194
f: 194,54
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t7_zfmisc_1
% SZS output start Proof for zfmisc_1__t7_zfmisc_1
1. ! [X0] : ! [X1] : r1_tarski(k1_tarski(X0),k2_tarski(X0,X1)) [input t7_zfmisc_1]
2. ~! [X0] : ! [X1] : r1_tarski(k1_tarski(X0),k2_tarski(X0,X1)) [negated conjecture 1]
5. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
6. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
8. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
22. ~! [X0,X1] : r1_tarski(k1_tarski(X0),k2_tarski(X0,X1)) [flattening 2]
23. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 8]
24. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 5]
27. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 6]
29. ? [X0,X1] : ~r1_tarski(k1_tarski(X0),k2_tarski(X0,X1)) [ennf transformation 22]
31. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 23]
35. ? [X0,X1] : ~r1_tarski(k1_tarski(X0),k2_tarski(X0,X1)) => ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) [choice axiom]
36. ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) [skolemisation 29,35]
37. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 31]
38. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 37]
39. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) [choice axiom]
40. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 38,39]
41. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 24]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 41]
43. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
44. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 42,43]
45. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 27]
46. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 45]
47. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 46]
48. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
49. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 47,48]
54. ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) [cnf transformation 36] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
61. r2_hidden(sK2(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
62. ~r2_hidden(sK2(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
63. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 44] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
70. r2_hidden(X4,X2) | X0 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 49] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
83. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 63] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
86. r2_hidden(X4,X2) | k2_tarski(X4,X1) != X2 [equality resolution 70] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
87. r2_hidden(X4,k2_tarski(X4,X1)) [equality resolution 86] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
90. 1 <=> r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) [avatar definition]
92. ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) <- (~1) [avatar component clause 90] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
93. ~1 [avatar split clause 54,90]
139. sK2(k1_tarski(X4),X5) = X4 | r1_tarski(k1_tarski(X4),X5) [resolution 61,83] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
175. ~r2_hidden(X2,X3) | r1_tarski(k1_tarski(X2),X3) | r1_tarski(k1_tarski(X2),X3) [superposition 62,139] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:6,thDist:-6}
178. r1_tarski(k1_tarski(X2),X3) | ~r2_hidden(X2,X3) [duplicate literal removal 175] {a:2,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
191. ~r2_hidden(sK0,k2_tarski(sK0,sK1)) <- (~1) [resolution 178,92] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
194. $false <- (~1) [subsumption resolution 191,87] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
195. 1 [avatar contradiction clause 194]
196. $false [avatar sat refutation 93,195] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:33,thDist:-33}
% SZS output end Proof for zfmisc_1__t7_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 108
% Active clauses: 38
% Passive clauses: 60
% Final active clauses: 37
% Final passive clauses: 18
% Main loop iterations started: 42
% 
% >>> Simplifying Inferences
% Duplicate literals: 6
% Fw subsumption resolutions: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 3
% Forward subsumptions: 31
% 
% >>> Generating Inferences
% Binary resolution: 44
% Unit resulting resolution: 1
% Forward superposition: 11
% Backward superposition: 10
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 8
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% backward demodulation index maintenance: 0.001 s
% resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.002 s ( own 0.0000 s ) 

