i: [1,30,0,6,2,1,0,0] t104_xboole_1
a: [3,46,0,3,1,30]
s: 46
a: [3,49,0,3,1,30]
s: 49
i: [1,29,0,6,2,1,0,0] t104_xboole_1
a: [3,55,0,3,1,29]
s: 55
i: [1,28,0,6,2,1,0,0] t104_xboole_1
a: [3,60,0,3,1,28]
s: 60
i: [1,37,0,3,1,0,0,1] irreflexivity_r2_xboole_0
s: 37
i: [1,38,0,3,1,0,0,2] reflexivity_r1_tarski
s: 38
i: [1,39,0,3,1,0,0,1] reflexivity_r3_xboole_0
s: 39
i: [1,42,0,3,1,0,0,1] d8_xboole_0
s: 42
i: [1,31,0,6,2,0,0,1] d8_xboole_0
s: 31
i: [1,35,0,6,2,0,0,2] d9_xboole_0
s: 35
i: [1,36,0,6,2,0,0,2] d9_xboole_0
s: 36
i: [1,40,0,6,2,0,0,1] antisymmetry_r2_xboole_0
s: 40
i: [1,41,0,6,2,0,0,1] symmetry_r3_xboole_0
s: 41
d: [2,65,1,3,1,1,40,41,49]
a: [3,70,1,3,1,65]
s: 70
i: [1,33,0,9,3,0,0,1] d8_xboole_0
s: 33
i: [1,34,0,9,3,0,0,2] d9_xboole_0
s: 34
d: [2,74,1,6,2,1,40,34,49]
a: [3,81,1,3,1,74]
s: 81
d: [2,88,2,0,0,3,52,46,55,81,33]
e: 88
d: [2,89,2,6,2,1,40,81,33]
d: [2,93,2,3,1,2,21,89,55]
d: [2,94,2,0,0,3,21,93,46]
e: 94
r: 46
r: 49
r: 70
r: 81
a: [3,45,0,3,1,30]
s: 45
r: 45
r: 60
s: 46
a: [3,50,0,3,1,30]
s: 50
a: [3,59,0,3,1,28]
s: 59
a: [3,69,1,3,1,65]
s: 69
a: [3,80,1,3,1,74]
s: 80
a: [3,85,1,3,1,74]
s: 85
d: [2,105,2,0,0,2,52,69,85,36]
e: 105
d: [2,106,2,0,0,2,52,50,85,35]
e: 106
d: [2,108,2,3,1,1,40,85,36]
d: [2,112,2,0,0,2,21,108,69]
e: 112
d: [2,109,2,3,1,1,40,85,35]
d: [2,114,2,0,0,2,21,109,50]
e: 114
r: 55
r: 59
r: 85
a: [3,54,0,3,1,29]
s: 54
r: 46
r: 50
r: 69
r: 80
d: [2,117,0,3,1,2,23,50,54]
d: [2,120,0,0,0,2,21,117,39]
e: 120
d: [2,118,1,3,1,2,23,69,54]
d: [2,122,1,0,0,2,21,118,39]
e: 122
d: [2,119,1,3,1,2,23,80,54]
d: [2,124,1,0,0,2,21,119,38]
e: 124
r: 54
s: 46
s: 49
s: 55
s: 60
s: 70
s: 80
s: 85
d: [2,133,2,0,0,3,52,60,55,85,33]
e: 133
d: [2,134,2,6,2,1,40,85,33]
d: [2,138,2,3,1,2,21,134,55]
d: [2,139,2,0,0,3,21,138,60]
e: 139
i: [1,27,0,12,4,1,0,0] t104_xboole_1
d: [2,90,2,3,1,1,40,81,36]
d: [2,100,1,3,1,1,40,45,31]
d: [2,103,1,3,1,1,40,59,31]
d: [2,126,1,6,2,1,40,49,34]
f: 139,126,124,114,103,100,90,93,27,28,29,30
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t104_xboole_1
% SZS output start Proof for xboole_1__t104_xboole_1
1. ! [X0] : ! [X1] : (~(~r2_xboole_0(X1,X0) & X0 != X1 & ~r2_xboole_0(X0,X1)) <=> r3_xboole_0(X0,X1)) [input t104_xboole_1]
2. ~! [X0] : ! [X1] : (~(~r2_xboole_0(X1,X0) & X0 != X1 & ~r2_xboole_0(X0,X1)) <=> r3_xboole_0(X0,X1)) [negated conjecture 1]
4. ! [X0] : ! [X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [input d8_xboole_0]
5. ! [X0] : ! [X1] : (r3_xboole_0(X0,X1) <=> (r1_tarski(X1,X0) | r1_tarski(X0,X1))) [input d9_xboole_0]
7. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
10. ~! [X0,X1] : (~(~r2_xboole_0(X1,X0) & X0 != X1 & ~r2_xboole_0(X0,X1)) <=> r3_xboole_0(X0,X1)) [flattening 2]
11. ! [X0,X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [flattening 4]
12. ! [X0,X1] : (r3_xboole_0(X0,X1) <=> (r1_tarski(X1,X0) | r1_tarski(X0,X1))) [flattening 5]
14. ! [X0] : r1_tarski(X0,X0) [rectify 7]
16. ? [X0,X1] : ((r2_xboole_0(X1,X0) | X0 = X1 | r2_xboole_0(X0,X1)) <~> r3_xboole_0(X0,X1)) [ennf transformation 10]
19. ? [X0,X1] : ((~r3_xboole_0(X0,X1) | (~r2_xboole_0(X1,X0) & X0 != X1 & ~r2_xboole_0(X0,X1))) & (r3_xboole_0(X0,X1) | (r2_xboole_0(X1,X0) | X0 = X1 | r2_xboole_0(X0,X1)))) [nnf transformation 16]
20. ? [X0,X1] : ((~r3_xboole_0(X0,X1) | (~r2_xboole_0(X1,X0) & X0 != X1 & ~r2_xboole_0(X0,X1))) & (r3_xboole_0(X0,X1) | r2_xboole_0(X1,X0) | X0 = X1 | r2_xboole_0(X0,X1))) [flattening 19]
21. ? [X0,X1] : ((~r3_xboole_0(X0,X1) | (~r2_xboole_0(X1,X0) & X0 != X1 & ~r2_xboole_0(X0,X1))) & (r3_xboole_0(X0,X1) | r2_xboole_0(X1,X0) | X0 = X1 | r2_xboole_0(X0,X1))) => ((~r3_xboole_0(sK0,sK1) | (~r2_xboole_0(sK1,sK0) & sK0 != sK1 & ~r2_xboole_0(sK0,sK1))) & (r3_xboole_0(sK0,sK1) | r2_xboole_0(sK1,sK0) | sK0 = sK1 | r2_xboole_0(sK0,sK1))) [choice axiom]
22. (~r3_xboole_0(sK0,sK1) | (~r2_xboole_0(sK1,sK0) & sK0 != sK1 & ~r2_xboole_0(sK0,sK1))) & (r3_xboole_0(sK0,sK1) | r2_xboole_0(sK1,sK0) | sK0 = sK1 | r2_xboole_0(sK0,sK1)) [skolemisation 20,21]
23. ! [X0,X1] : ((r2_xboole_0(X0,X1) | (X0 = X1 | ~r1_tarski(X0,X1))) & ((X0 != X1 & r1_tarski(X0,X1)) | ~r2_xboole_0(X0,X1))) [nnf transformation 11]
24. ! [X0,X1] : ((r2_xboole_0(X0,X1) | X0 = X1 | ~r1_tarski(X0,X1)) & ((X0 != X1 & r1_tarski(X0,X1)) | ~r2_xboole_0(X0,X1))) [flattening 23]
25. ! [X0,X1] : ((r3_xboole_0(X0,X1) | (~r1_tarski(X1,X0) & ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) | r1_tarski(X0,X1)) | ~r3_xboole_0(X0,X1))) [nnf transformation 12]
26. ! [X0,X1] : ((r3_xboole_0(X0,X1) | (~r1_tarski(X1,X0) & ~r1_tarski(X0,X1))) & (r1_tarski(X1,X0) | r1_tarski(X0,X1) | ~r3_xboole_0(X0,X1))) [flattening 25]
27. r3_xboole_0(sK0,sK1) | r2_xboole_0(sK1,sK0) | sK0 = sK1 | r2_xboole_0(sK0,sK1) [cnf transformation 22] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. ~r3_xboole_0(sK0,sK1) | ~r2_xboole_0(sK0,sK1) [cnf transformation 22] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~r3_xboole_0(sK0,sK1) | sK0 != sK1 [cnf transformation 22] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. ~r3_xboole_0(sK0,sK1) | ~r2_xboole_0(sK1,sK0) [cnf transformation 22] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~r2_xboole_0(X0,X1) | r1_tarski(X0,X1) [cnf transformation 24] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. ~r1_tarski(X0,X1) | X0 = X1 | r2_xboole_0(X0,X1) [cnf transformation 24] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. ~r3_xboole_0(X0,X1) | r1_tarski(X0,X1) | r1_tarski(X1,X0) [cnf transformation 26] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
35. ~r1_tarski(X0,X1) | r3_xboole_0(X0,X1) [cnf transformation 26] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
36. ~r1_tarski(X1,X0) | r3_xboole_0(X0,X1) [cnf transformation 26] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
38. r1_tarski(X0,X0) [cnf transformation 14] {a:0,w:3,wCS:1200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
44. 1 <=> r2_xboole_0(sK1,sK0) [avatar definition]
45. r2_xboole_0(sK1,sK0) <- (1) [avatar component clause 44] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
48. 2 <=> r3_xboole_0(sK0,sK1) [avatar definition]
49. r3_xboole_0(sK0,sK1) <- (2) [avatar component clause 48] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
50. ~r3_xboole_0(sK0,sK1) <- (~2) [avatar component clause 48] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
51. ~1 | ~2 [avatar split clause 30,48,44]
53. 3 <=> sK0 = sK1 [avatar definition]
54. sK0 = sK1 <- (3) [avatar component clause 53] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
55. sK0 != sK1 <- (~3) [avatar component clause 53] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. ~3 | ~2 [avatar split clause 29,48,53]
58. 4 <=> r2_xboole_0(sK0,sK1) [avatar definition]
59. r2_xboole_0(sK0,sK1) <- (4) [avatar component clause 58] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. ~r2_xboole_0(sK0,sK1) <- (~4) [avatar component clause 58] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
61. ~4 | ~2 [avatar split clause 28,48,58]
62. 4 | 3 | 1 | 2 [avatar split clause 27,48,44,53,58]
79. 6 <=> r1_tarski(sK1,sK0) [avatar definition]
80. ~r1_tarski(sK1,sK0) <- (~6) [avatar component clause 79] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
81. r1_tarski(sK1,sK0) <- (6) [avatar component clause 79] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
83. 7 <=> r1_tarski(sK0,sK1) [avatar definition]
85. r1_tarski(sK0,sK1) <- (7) [avatar component clause 83] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. sK0 = sK1 | r2_xboole_0(sK1,sK0) <- (6) [resolution 81,33] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
90. r3_xboole_0(sK0,sK1) <- (6) [resolution 81,36] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
93. r2_xboole_0(sK1,sK0) <- (~3, 6) [subsumption resolution 89,55] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
96. 1 | 3 | ~6 [avatar split clause 93,79,53,44]
97. 2 | ~6 [avatar split clause 90,79,48]
100. r1_tarski(sK1,sK0) <- (1) [resolution 45,31] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
101. 6 | ~1 [avatar split clause 100,44,79]
103. r1_tarski(sK0,sK1) <- (4) [resolution 59,31] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
104. 7 | ~4 [avatar split clause 103,58,83]
109. r3_xboole_0(sK0,sK1) <- (7) [resolution 85,35] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
114. $false <- (~2, 7) [subsumption resolution 109,50] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
115. 2 | ~7 [avatar contradiction clause 114]
119. ~r1_tarski(sK0,sK0) <- (3, ~6) [backward demodulation 80,54] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. $false <- (3, ~6) [subsumption resolution 119,38] {a:1,w:0,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
125. ~3 | 6 [avatar contradiction clause 124]
126. r1_tarski(sK0,sK1) | r1_tarski(sK1,sK0) <- (2) [resolution 49,34] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
128. 6 | 7 | ~2 [avatar split clause 126,48,83,79]
134. sK0 = sK1 | r2_xboole_0(sK0,sK1) <- (7) [resolution 85,33] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
138. r2_xboole_0(sK0,sK1) <- (~3, 7) [subsumption resolution 134,55] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
139. $false <- (~3, ~4, 7) [subsumption resolution 138,60] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
140. 3 | 4 | ~7 [avatar contradiction clause 139]
141. $false [avatar sat refutation 51,56,61,62,96,97,101,104,115,125,128,140] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:136,thDist:-136}
% SZS output end Proof for xboole_1__t104_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 83
% Active clauses: 31
% Passive clauses: 32
% Final active clauses: 17
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 9
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Equational tautologies: 1
% Forward subsumptions: 15
% 
% >>> Generating Inferences
% Binary resolution: 26
% Unit resulting resolution: 4
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 18
% Sat splitting refutations: 11
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 28
% SAT solver binary clauses: 16
% 
% Memory used [KB]: 9338
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

