i: [1,45,0,2,1,1,0,0] t74_funct_4
a: [3,74,0,2,1,45]
s: 74
i: [1,44,0,2,1,1,0,0] t74_funct_4
a: [3,79,0,2,1,44]
s: 79
i: [1,43,0,2,1,1,0,0] t74_funct_4
a: [3,84,0,2,1,43]
s: 84
i: [1,42,0,2,1,1,0,0] t74_funct_4
a: [3,89,0,2,1,42]
s: 89
i: [1,41,0,2,1,1,0,0] t74_funct_4
a: [3,94,0,2,1,41]
s: 94
i: [1,40,0,2,1,1,0,0] t74_funct_4
a: [3,99,0,2,1,40]
s: 99
i: [1,46,0,5,1,1,0,0] t74_funct_4
a: [3,69,0,5,1,46]
s: 69
i: [1,60,0,2,1,0,0,1] rc1_funct_1
a: [3,111,0,2,1,60]
s: 111
i: [1,59,0,2,1,0,0,1] rc1_funct_1
a: [3,116,0,2,1,59]
s: 116
i: [1,47,0,9,1,1,0,0] t74_funct_4
a: [3,64,0,9,1,47]
s: 64
i: [1,55,0,3,1,0,0,2] reflexivity_r1_tarski
s: 55
i: [1,54,0,4,1,0,0,3] existence_m1_subset_1
s: 54
i: [1,52,0,7,2,0,0,2] t3_subset
s: 52
d: [2,118,1,5,1,0,40,52,54]
s: 118
i: [1,53,0,7,2,0,0,2] t3_subset
s: 53
i: [1,48,0,8,3,0,0,3] cc2_relat_1
s: 48
d: [2,120,1,6,2,0,40,48,54]
s: 120
i: [1,56,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,106,0,9,3,56]
s: 106
d: [2,121,1,7,3,0,40,48,53]
s: 121
i: [1,49,0,10,4,0,0,3] cc3_funct_1
s: 49
i: [1,57,0,12,5,0,0,1] dt_k1_funct_4
s: 57
d: [2,124,1,8,3,0,40,49,54]
s: 124
i: [1,58,0,12,5,0,0,1] dt_k1_funct_4
s: 58
d: [2,125,1,9,4,0,40,49,53]
s: 125
i: [1,50,0,18,6,0,0,1] t19_funct_4
s: 50
i: [1,51,0,23,7,0,0,1] t14_funct_4
s: 51
d: [2,143,1,21,7,1,44,64,51]
d: [2,157,1,19,6,2,21,143,74]
d: [2,158,1,17,5,3,21,157,79]
d: [2,159,1,15,4,4,21,158,84]
d: [2,160,1,13,3,5,21,159,89]
d: [2,161,1,11,2,6,21,160,94]
d: [2,162,1,9,1,7,21,161,99]
a: [3,166,1,9,1,162]
s: 166
d: [2,133,1,18,6,1,44,50,69]
d: [2,136,1,16,5,2,21,133,74]
d: [2,137,1,14,4,3,21,136,79]
s: 137
d: [2,174,2,9,3,3,40,137,55]
d: [2,176,2,7,2,4,21,174,84]
d: [2,177,2,5,1,5,21,176,89]
a: [3,181,2,5,1,177]
s: 181
d: [2,134,1,18,6,1,44,50,69]
d: [2,138,1,16,5,2,21,134,74]
d: [2,139,1,14,4,3,21,138,79]
s: 139
d: [2,190,2,9,3,3,40,139,55]
d: [2,192,2,7,2,4,21,190,84]
d: [2,193,2,5,1,5,21,192,89]
f: 193,162,40,41,42,43,44,45,46,47
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t74_funct_4
% SZS output start Proof for funct_4__t74_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (k9_xtuple_0(X1) = k9_xtuple_0(X2) => k1_funct_4(k1_funct_4(X0,X1),X2) = k1_funct_4(X0,X2))))) [input t74_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (k9_xtuple_0(X1) = k9_xtuple_0(X2) => k1_funct_4(k1_funct_4(X0,X1),X2) = k1_funct_4(X0,X2))))) [negated conjecture 1]
12. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
13. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => k1_funct_4(k1_funct_4(X0,X1),X2) = k1_funct_4(X0,k1_funct_4(X1,X2))))) [input t14_funct_4]
14. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k1_funct_4(X0,X1) = X1))) [input t19_funct_4]
17. ! [X0] : r1_tarski(X0,X0) [rectify 12]
18. ? [X0] : (? [X1] : (? [X2] : ((k1_funct_4(k1_funct_4(X0,X1),X2) != k1_funct_4(X0,X2) & k9_xtuple_0(X1) = k9_xtuple_0(X2)) & (v1_funct_1(X2) & v1_relat_1(X2))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
19. ? [X0] : (? [X1] : (? [X2] : (k1_funct_4(k1_funct_4(X0,X1),X2) != k1_funct_4(X0,X2) & k9_xtuple_0(X1) = k9_xtuple_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 18]
23. ! [X0] : (! [X1] : ((k1_funct_4(X0,X1) = X1 | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 14]
24. ! [X0] : (! [X1] : (k1_funct_4(X0,X1) = X1 | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 23]
25. ! [X0] : (! [X1] : (! [X2] : (k1_funct_4(k1_funct_4(X0,X1),X2) = k1_funct_4(X0,k1_funct_4(X1,X2)) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 13]
26. ! [X0] : (! [X1] : (! [X2] : (k1_funct_4(k1_funct_4(X0,X1),X2) = k1_funct_4(X0,k1_funct_4(X1,X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 25]
31. ? [X0] : (? [X1] : (? [X2] : (k1_funct_4(k1_funct_4(X0,X1),X2) != k1_funct_4(X0,X2) & k9_xtuple_0(X1) = k9_xtuple_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (k1_funct_4(k1_funct_4(sK0,X1),X2) != k1_funct_4(sK0,X2) & k9_xtuple_0(X1) = k9_xtuple_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
32. ? [X1] : (? [X2] : (k1_funct_4(k1_funct_4(sK0,X1),X2) != k1_funct_4(sK0,X2) & k9_xtuple_0(X1) = k9_xtuple_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) => (? [X2] : (k1_funct_4(sK0,X2) != k1_funct_4(k1_funct_4(sK0,sK1),X2) & k9_xtuple_0(X2) = k9_xtuple_0(sK1) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
33. ? [X2] : (k1_funct_4(sK0,X2) != k1_funct_4(k1_funct_4(sK0,sK1),X2) & k9_xtuple_0(X2) = k9_xtuple_0(sK1) & v1_funct_1(X2) & v1_relat_1(X2)) => (k1_funct_4(k1_funct_4(sK0,sK1),sK2) != k1_funct_4(sK0,sK2) & k9_xtuple_0(sK1) = k9_xtuple_0(sK2) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
34. ((k1_funct_4(k1_funct_4(sK0,sK1),sK2) != k1_funct_4(sK0,sK2) & k9_xtuple_0(sK1) = k9_xtuple_0(sK2) & v1_funct_1(sK2) & v1_relat_1(sK2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 19,33,32,31]
40. v1_relat_1(sK0) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
41. v1_funct_1(sK0) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
42. v1_relat_1(sK1) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
43. v1_funct_1(sK1) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
44. v1_relat_1(sK2) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
45. v1_funct_1(sK2) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
46. k9_xtuple_0(sK1) = k9_xtuple_0(sK2) [cnf transformation 34] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
47. k1_funct_4(k1_funct_4(sK0,sK1),sK2) != k1_funct_4(sK0,sK2) [cnf transformation 34] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
50. ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | k1_funct_4(X0,X1) = X1 | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k1_funct_4(k1_funct_4(X0,X1),X2) = k1_funct_4(X0,k1_funct_4(X1,X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. r1_tarski(X0,X0) [cnf transformation 17] {a:0,w:3,wCS:1200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
62. 1 <=> k1_funct_4(k1_funct_4(sK0,sK1),sK2) = k1_funct_4(sK0,sK2) [avatar definition]
64. k1_funct_4(k1_funct_4(sK0,sK1),sK2) != k1_funct_4(sK0,sK2) <- (~1) [avatar component clause 62] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
65. ~1 [avatar split clause 47,62]
67. 2 <=> k9_xtuple_0(sK1) = k9_xtuple_0(sK2) [avatar definition]
69. k9_xtuple_0(sK1) = k9_xtuple_0(sK2) <- (2) [avatar component clause 67] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
70. 2 [avatar split clause 46,67]
72. 3 <=> v1_funct_1(sK2) [avatar definition]
74. v1_funct_1(sK2) <- (3) [avatar component clause 72] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
75. 3 [avatar split clause 45,72]
77. 4 <=> v1_relat_1(sK2) [avatar definition]
79. v1_relat_1(sK2) <- (4) [avatar component clause 77] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
80. 4 [avatar split clause 44,77]
82. 5 <=> v1_funct_1(sK1) [avatar definition]
84. v1_funct_1(sK1) <- (5) [avatar component clause 82] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
85. 5 [avatar split clause 43,82]
87. 6 <=> v1_relat_1(sK1) [avatar definition]
89. v1_relat_1(sK1) <- (6) [avatar component clause 87] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
90. 6 [avatar split clause 42,87]
92. 7 <=> v1_funct_1(sK0) [avatar definition]
94. v1_funct_1(sK0) <- (7) [avatar component clause 92] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
95. 7 [avatar split clause 41,92]
97. 8 <=> v1_relat_1(sK0) [avatar definition]
99. v1_relat_1(sK0) <- (8) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
100. 8 [avatar split clause 40,97]
134. ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(sK1)) | sK2 = k1_funct_4(X0,sK2) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2) [superposition 50,69] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
138. ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(sK1)) | sK2 = k1_funct_4(X0,sK2) | ~v1_relat_1(sK2) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2, 3) [subsumption resolution 134,74] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
139. ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(sK1)) | sK2 = k1_funct_4(X0,sK2) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2, 3, 4) [subsumption resolution 138,79] {a:1,w:14,wCS:1400,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
143. k1_funct_4(sK0,sK2) != k1_funct_4(sK0,k1_funct_4(sK1,sK2)) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [superposition 64,51] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
157. k1_funct_4(sK0,sK2) != k1_funct_4(sK0,k1_funct_4(sK1,sK2)) | ~v1_relat_1(sK2) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 3) [subsumption resolution 143,74] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
158. k1_funct_4(sK0,sK2) != k1_funct_4(sK0,k1_funct_4(sK1,sK2)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 3, 4) [subsumption resolution 157,79] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
159. k1_funct_4(sK0,sK2) != k1_funct_4(sK0,k1_funct_4(sK1,sK2)) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 3, 4, 5) [subsumption resolution 158,84] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
160. k1_funct_4(sK0,sK2) != k1_funct_4(sK0,k1_funct_4(sK1,sK2)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 3, 4, 5, 6) [subsumption resolution 159,89] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
161. k1_funct_4(sK0,sK2) != k1_funct_4(sK0,k1_funct_4(sK1,sK2)) | ~v1_relat_1(sK0) <- (~1, 3, 4, 5, 6, 7) [subsumption resolution 160,94] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
162. k1_funct_4(sK0,sK2) != k1_funct_4(sK0,k1_funct_4(sK1,sK2)) <- (~1, 3, 4, 5, 6, 7, 8) [subsumption resolution 161,99] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
164. 13 <=> k1_funct_4(sK0,sK2) = k1_funct_4(sK0,k1_funct_4(sK1,sK2)) [avatar definition]
167. ~13 | 1 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 162,97,92,87,82,77,72,62,164]
190. sK2 = k1_funct_4(sK1,sK2) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (2, 3, 4) [resolution 139,55] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
192. sK2 = k1_funct_4(sK1,sK2) | ~v1_relat_1(sK1) <- (2, 3, 4, 5) [subsumption resolution 190,84] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
193. sK2 = k1_funct_4(sK1,sK2) <- (2, 3, 4, 5, 6) [subsumption resolution 192,89] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
195. 15 <=> sK2 = k1_funct_4(sK1,sK2) [avatar definition]
198. 15 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar split clause 193,87,82,77,72,67,195]
199. sK2 != k1_funct_4(sK1,sK2) | k1_funct_4(sK0,sK2) = k1_funct_4(sK0,k1_funct_4(sK1,sK2)) [theory tautology sat conflict] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
200. $false [avatar sat refutation 65,70,75,80,85,90,95,100,167,198,199] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:87,thDist:-87}
% SZS output end Proof for funct_4__t74_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 100
% Active clauses: 30
% Passive clauses: 37
% Final active clauses: 30
% Final passive clauses: 7
% Main loop iterations started: 30
% 
% >>> Simplifying Inferences
% Duplicate literals: 24
% Fw subsumption resolutions: 28
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 13
% Forward superposition: 8
% Backward superposition: 7
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 10
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

