i: [1,79,0,2,1,1,0,0] t22_xtuple_0
a: [3,147,0,2,1,79]
s: 147
i: [1,78,0,2,1,1,0,0] t22_xtuple_0
a: [3,152,0,2,1,78]
s: 152
i: [1,84,0,3,1,1,0,0] t22_xtuple_0
a: [3,122,0,3,1,84]
s: 122
i: [1,83,0,5,1,1,0,0] t22_xtuple_0
a: [3,127,0,5,1,83]
s: 127
i: [1,82,0,5,1,1,0,0] t22_xtuple_0
a: [3,132,0,5,1,82]
s: 132
i: [1,81,0,5,1,1,0,0] t22_xtuple_0
a: [3,137,0,5,1,81]
s: 137
i: [1,80,0,5,1,1,0,0] t22_xtuple_0
a: [3,142,0,5,1,80]
s: 142
i: [1,85,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,157,0,2,1,85]
s: 157
i: [1,86,0,2,1,0,0,255] fc1_xboole_0
a: [3,162,0,2,1,86]
s: 162
i: [1,114,0,2,1,0,0,255] rc2_xboole_0
a: [3,172,0,2,1,114]
s: 172
i: [1,115,0,2,1,0,0,255] rc2_xtuple_0
a: [3,177,0,2,1,115]
s: 177
i: [1,116,0,2,1,0,0,255] rc1_xtuple_0
a: [3,182,0,2,1,116]
s: 182
i: [1,117,0,2,1,0,0,255] rc1_xboole_0
a: [3,187,0,2,1,117]
s: 187
i: [1,118,0,2,1,0,0,1] rc3_xtuple_0
a: [3,192,0,2,1,118]
s: 192
i: [1,88,0,3,1,0,0,255] fc2_xboole_0
s: 88
i: [1,87,0,3,1,0,0,255] d2_xboole_0
a: [3,167,0,3,1,87]
s: 167
r: 157
i: [1,92,0,4,2,0,0,255] cc1_xtuple_0
s: 92
i: [1,95,0,4,2,0,0,255] cc2_xtuple_0
s: 95
d: [2,201,1,2,1,1,40,95,147]
a: [3,207,1,2,1,201]
s: 207
i: [1,101,0,4,1,0,0,255] fc3_xboole_0
s: 101
i: [1,102,0,4,1,0,0,255] fc1_xtuple_0
s: 102
d: [2,202,1,2,1,1,40,95,152]
a: [3,212,1,2,1,202]
s: 212
i: [1,94,0,5,2,0,0,255] t6_boole
s: 94
d: [2,219,2,2,1,1,40,207,92]
a: [3,223,2,2,1,219]
s: 223
i: [1,106,0,5,1,0,0,255] fc2_xtuple_0
s: 106
i: [1,89,0,6,1,0,0,1] d7_xtuple_0
s: 89
d: [2,225,2,2,1,1,40,212,92]
a: [3,229,2,2,1,225]
s: 229
i: [1,103,0,6,1,0,0,2] rd1_xtuple_0
s: 103
d: [2,195,1,2,1,1,40,92,177]
a: [3,199,1,2,1,195]
s: 199
i: [1,104,0,6,1,0,0,3] rd2_xtuple_0
s: 104
i: [1,109,0,6,1,0,0,1] fc3_xtuple_0
s: 109
d: [2,203,1,2,1,1,40,95,192]
a: [3,217,1,2,1,203]
s: 217
i: [1,90,0,7,1,0,0,1] d10_xtuple_0
s: 90
d: [2,236,2,2,1,1,40,217,92]
a: [3,240,2,2,1,236]
s: 240
i: [1,91,0,7,1,0,0,1] d11_xtuple_0
s: 91
i: [1,98,0,7,3,0,0,255] t8_boole
s: 98
d: [2,233,1,5,1,0,40,106,92]
s: 233
i: [1,105,0,7,1,0,0,255] commutativity_k2_tarski
s: 105
d: [2,235,1,6,1,0,40,109,95]
s: 235
i: [1,107,0,7,1,0,0,1] rd5_xtuple_0
s: 107
i: [1,108,0,7,1,0,0,2] rd6_xtuple_0
s: 108
d: [2,249,1,6,1,0,44,91,89]
s: 249
i: [1,110,0,8,1,0,0,1] rd8_xtuple_0
s: 110
d: [2,250,1,6,1,0,44,89,91]
s: 250
i: [1,111,0,8,1,0,0,1] rd9_xtuple_0
s: 111
i: [1,112,0,8,1,0,0,1] rd10_xtuple_0
s: 112
d: [2,255,2,6,1,0,40,235,92]
s: 255
i: [1,113,0,8,1,0,0,2] rd11_xtuple_0
s: 113
d: [2,231,1,7,3,0,44,94,94]
s: 231
i: [1,93,0,9,2,0,0,255] rd3_xtuple_0
s: 93
i: [1,97,0,10,1,0,0,255] d5_tarski
s: 97
d: [2,265,1,9,2,1,44,93,127]
d: [2,274,1,7,1,2,21,265,223]
a: [3,278,1,7,1,274]
s: 278
i: [1,99,0,10,1,0,0,2] d4_xtuple_0
s: 99
d: [2,291,2,10,1,1,44,99,278]
s: 291
i: [1,100,0,12,1,0,0,2] d8_xtuple_0
s: 100
i: [1,96,0,13,2,0,0,1] rd12_xtuple_0
s: 96
d: [2,299,3,4,1,1,44,106,291]
s: 299
d: [2,306,1,13,2,1,44,96,142]
d: [2,323,1,11,1,2,21,306,147]
d: [2,324,1,11,1,3,22,323,127]
d: [2,325,1,11,1,4,22,324,132]
d: [2,326,1,11,1,5,22,325,137]
a: [3,330,1,11,1,326]
s: 330
d: [2,365,2,5,2,1,44,330,96]
d: [2,374,2,2,1,2,21,365,122]
d: [2,375,2,0,0,3,21,374,152]
e: 375
d: [2,366,2,5,2,1,44,96,330]
d: [2,377,2,2,1,2,21,366,122]
d: [2,378,2,0,0,3,21,377,152]
e: 378
d: [2,319,1,13,2,1,44,96,127]
d: [2,359,1,11,1,2,21,319,147]
d: [2,360,1,11,1,3,22,359,132]
d: [2,361,1,11,1,4,22,360,137]
d: [2,362,1,11,1,5,22,361,142]
f: 378,362,78,79,80,81,82,83,84
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t22_xtuple_0
% SZS output start Proof for xtuple_0__t22_xtuple_0
1. ! [X0] : (v3_xtuple_0(X0) => ! [X1] : (v3_xtuple_0(X1) => ((k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k8_xtuple_0(X0) = k8_xtuple_0(X1) & k7_xtuple_0(X0) = k7_xtuple_0(X1)) => X0 = X1))) [input t22_xtuple_0]
2. ~! [X0] : (v3_xtuple_0(X0) => ! [X1] : (v3_xtuple_0(X1) => ((k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k8_xtuple_0(X0) = k8_xtuple_0(X1) & k7_xtuple_0(X0) = k7_xtuple_0(X1)) => X0 = X1))) [negated conjecture 1]
39. ! [X0] : (v3_xtuple_0(X0) => k6_xtuple_0(k7_xtuple_0(X0),k8_xtuple_0(X0),k5_xtuple_0(X0),k2_xtuple_0(X0)) = X0) [input rd12_xtuple_0]
57. ? [X0] : (? [X1] : ((X0 != X1 & (k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k8_xtuple_0(X0) = k8_xtuple_0(X1) & k7_xtuple_0(X0) = k7_xtuple_0(X1))) & v3_xtuple_0(X1)) & v3_xtuple_0(X0)) [ennf transformation 2]
58. ? [X0] : (? [X1] : (X0 != X1 & k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k8_xtuple_0(X0) = k8_xtuple_0(X1) & k7_xtuple_0(X0) = k7_xtuple_0(X1) & v3_xtuple_0(X1)) & v3_xtuple_0(X0)) [flattening 57]
63. ! [X0] : (k6_xtuple_0(k7_xtuple_0(X0),k8_xtuple_0(X0),k5_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v3_xtuple_0(X0)) [ennf transformation 39]
65. ? [X0] : (? [X1] : (X0 != X1 & k2_xtuple_0(X0) = k2_xtuple_0(X1) & k5_xtuple_0(X0) = k5_xtuple_0(X1) & k8_xtuple_0(X0) = k8_xtuple_0(X1) & k7_xtuple_0(X0) = k7_xtuple_0(X1) & v3_xtuple_0(X1)) & v3_xtuple_0(X0)) => (? [X1] : (sK0 != X1 & k2_xtuple_0(X1) = k2_xtuple_0(sK0) & k5_xtuple_0(X1) = k5_xtuple_0(sK0) & k8_xtuple_0(X1) = k8_xtuple_0(sK0) & k7_xtuple_0(X1) = k7_xtuple_0(sK0) & v3_xtuple_0(X1)) & v3_xtuple_0(sK0)) [choice axiom]
66. ? [X1] : (sK0 != X1 & k2_xtuple_0(X1) = k2_xtuple_0(sK0) & k5_xtuple_0(X1) = k5_xtuple_0(sK0) & k8_xtuple_0(X1) = k8_xtuple_0(sK0) & k7_xtuple_0(X1) = k7_xtuple_0(sK0) & v3_xtuple_0(X1)) => (sK0 != sK1 & k2_xtuple_0(sK0) = k2_xtuple_0(sK1) & k5_xtuple_0(sK0) = k5_xtuple_0(sK1) & k8_xtuple_0(sK0) = k8_xtuple_0(sK1) & k7_xtuple_0(sK0) = k7_xtuple_0(sK1) & v3_xtuple_0(sK1)) [choice axiom]
67. (sK0 != sK1 & k2_xtuple_0(sK0) = k2_xtuple_0(sK1) & k5_xtuple_0(sK0) = k5_xtuple_0(sK1) & k8_xtuple_0(sK0) = k8_xtuple_0(sK1) & k7_xtuple_0(sK0) = k7_xtuple_0(sK1) & v3_xtuple_0(sK1)) & v3_xtuple_0(sK0) [skolemisation 58,66,65]
78. v3_xtuple_0(sK0) [cnf transformation 67] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
79. v3_xtuple_0(sK1) [cnf transformation 67] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
80. k7_xtuple_0(sK0) = k7_xtuple_0(sK1) [cnf transformation 67] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
81. k8_xtuple_0(sK0) = k8_xtuple_0(sK1) [cnf transformation 67] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
82. k5_xtuple_0(sK0) = k5_xtuple_0(sK1) [cnf transformation 67] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
83. k2_xtuple_0(sK0) = k2_xtuple_0(sK1) [cnf transformation 67] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
84. sK0 != sK1 [cnf transformation 67] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
96. k6_xtuple_0(k7_xtuple_0(X0),k8_xtuple_0(X0),k5_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v3_xtuple_0(X0) [cnf transformation 63] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. 1 <=> sK0 = sK1 [avatar definition]
122. sK0 != sK1 <- (~1) [avatar component clause 120] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
123. ~1 [avatar split clause 84,120]
125. 2 <=> k2_xtuple_0(sK0) = k2_xtuple_0(sK1) [avatar definition]
127. k2_xtuple_0(sK0) = k2_xtuple_0(sK1) <- (2) [avatar component clause 125] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. 2 [avatar split clause 83,125]
130. 3 <=> k5_xtuple_0(sK0) = k5_xtuple_0(sK1) [avatar definition]
132. k5_xtuple_0(sK0) = k5_xtuple_0(sK1) <- (3) [avatar component clause 130] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
133. 3 [avatar split clause 82,130]
135. 4 <=> k8_xtuple_0(sK0) = k8_xtuple_0(sK1) [avatar definition]
137. k8_xtuple_0(sK0) = k8_xtuple_0(sK1) <- (4) [avatar component clause 135] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
138. 4 [avatar split clause 81,135]
140. 5 <=> k7_xtuple_0(sK0) = k7_xtuple_0(sK1) [avatar definition]
142. k7_xtuple_0(sK0) = k7_xtuple_0(sK1) <- (5) [avatar component clause 140] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
143. 5 [avatar split clause 80,140]
145. 6 <=> v3_xtuple_0(sK1) [avatar definition]
147. v3_xtuple_0(sK1) <- (6) [avatar component clause 145] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
148. 6 [avatar split clause 79,145]
150. 7 <=> v3_xtuple_0(sK0) [avatar definition]
152. v3_xtuple_0(sK0) <- (7) [avatar component clause 150] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
153. 7 [avatar split clause 78,150]
319. sK1 = k6_xtuple_0(k7_xtuple_0(sK1),k8_xtuple_0(sK1),k5_xtuple_0(sK1),k2_xtuple_0(sK0)) | ~v3_xtuple_0(sK1) <- (2) [superposition 96,127] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
328. 24 <=> sK1 = k6_xtuple_0(k7_xtuple_0(sK0),k8_xtuple_0(sK0),k5_xtuple_0(sK0),k2_xtuple_0(sK0)) [avatar definition]
330. sK1 = k6_xtuple_0(k7_xtuple_0(sK0),k8_xtuple_0(sK0),k5_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (24) [avatar component clause 328] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
359. sK1 = k6_xtuple_0(k7_xtuple_0(sK1),k8_xtuple_0(sK1),k5_xtuple_0(sK1),k2_xtuple_0(sK0)) <- (2, 6) [subsumption resolution 319,147] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
360. sK1 = k6_xtuple_0(k7_xtuple_0(sK1),k8_xtuple_0(sK1),k5_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (2, 3, 6) [forward demodulation 359,132] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
361. sK1 = k6_xtuple_0(k7_xtuple_0(sK1),k8_xtuple_0(sK0),k5_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (2, 3, 4, 6) [forward demodulation 360,137] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
362. sK1 = k6_xtuple_0(k7_xtuple_0(sK0),k8_xtuple_0(sK0),k5_xtuple_0(sK0),k2_xtuple_0(sK0)) <- (2, 3, 4, 5, 6) [forward demodulation 361,142] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
363. 24 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar split clause 362,145,140,135,130,125,328]
366. sK0 = sK1 | ~v3_xtuple_0(sK0) <- (24) [superposition 96,330] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
377. ~v3_xtuple_0(sK0) <- (~1, 24) [subsumption resolution 366,122] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
378. $false <- (~1, 7, 24) [subsumption resolution 377,152] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
379. 1 | ~7 | ~24 [avatar contradiction clause 378]
380. $false [avatar sat refutation 123,128,133,138,143,148,153,363,379] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:148,thDist:-148}
% SZS output end Proof for xtuple_0__t22_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 48
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 201
% Active clauses: 56
% Passive clauses: 95
% Final active clauses: 55
% Final passive clauses: 37
% Main loop iterations started: 58
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 17
% Fw demodulations: 31
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 22
% Fw demodulations to eq. taut.: 6
% 
% >>> Generating Inferences
% Binary resolution: 11
% Forward superposition: 37
% Backward superposition: 34
% Self superposition: 5
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 31
% SAT solver unit clauses: 15
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9594
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% sat proof minimization: 0.001 s
% unit clause index maintenance: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

