i: [1,189,0,2,1,1,0,0] t13_toler_1
a: [3,314,0,2,1,189]
s: 314
i: [1,188,0,2,1,1,0,0] t13_toler_1
a: [3,319,0,2,1,188]
s: 319
i: [1,194,0,3,1,1,0,0] t13_toler_1
a: [3,289,0,3,1,194]
s: 289
i: [1,190,0,3,1,1,0,0] t13_toler_1
a: [3,309,0,3,1,190]
s: 309
i: [1,193,0,4,1,1,0,0] t13_toler_1
a: [3,294,0,4,1,193]
s: 294
i: [1,192,0,4,1,1,0,0] t13_toler_1
a: [3,299,0,4,1,192]
s: 299
i: [1,191,0,6,1,1,0,0] t13_toler_1
a: [3,304,0,6,1,191]
s: 304
i: [1,195,0,2,1,0,0,1] fc1_xboole_0
a: [3,324,0,2,1,195]
s: 324
i: [1,281,0,2,1,0,0,2] rc2_xboole_0
a: [3,329,0,2,1,281]
s: 329
i: [1,282,0,2,1,0,0,2] rc1_xboole_0
a: [3,334,0,2,1,282]
s: 334
i: [1,284,0,2,1,0,0,2] rc1_relat_1
a: [3,339,0,2,1,284]
s: 339
i: [1,283,0,2,1,0,0,2] rc1_relat_1
a: [3,344,0,2,1,283]
s: 344
i: [1,285,0,2,1,0,0,255] rc2_relat_1
a: [3,349,0,2,1,285]
s: 349
i: [1,196,0,3,1,0,0,255] fc2_xboole_0
s: 196
i: [1,197,0,3,1,0,0,1] fc1_subset_1
s: 197
i: [1,242,0,3,1,0,0,1] rc2_subset_1
s: 242
i: [1,246,0,3,1,0,0,255] rc2_partfun1
s: 246
i: [1,249,0,3,1,0,0,255] rc2_partfun1
s: 249
i: [1,250,0,3,1,0,0,255] rc2_partfun1
s: 250
i: [1,251,0,3,1,0,0,255] rc2_partfun1
s: 251
i: [1,253,0,3,1,0,0,255] reflexivity_r1_tarski
s: 253
i: [1,206,0,4,2,0,0,2] cc1_relat_1
s: 206
i: [1,240,0,4,1,0,0,1] existence_m1_subset_1
s: 240
d: [2,351,1,2,1,1,40,206,324]
a: [3,357,1,2,1,351]
s: 357
i: [1,244,0,4,1,0,0,255] rc3_subset_1
s: 244
i: [1,247,0,4,1,0,0,255] rc2_partfun1
s: 247
d: [2,352,1,2,1,1,40,206,334]
a: [3,362,1,2,1,352]
s: 362
i: [1,248,0,4,1,0,0,255] rc2_partfun1
s: 248
d: [2,353,1,3,1,0,40,206,242]
s: 353
i: [1,252,0,4,1,0,0,255] rc2_partfun1
s: 252
i: [1,254,0,4,1,0,0,255] fc3_xboole_0
s: 254
i: [1,255,0,4,1,0,0,1] fc6_relat_1
s: 255
i: [1,272,0,4,1,0,0,255] rc3_relat_1
s: 272
i: [1,276,0,4,1,0,0,255] rc1_relset_1
s: 276
i: [1,277,0,4,1,0,0,255] rc1_relset_1
s: 277
i: [1,201,0,5,2,0,0,1] rc1_subset_1
s: 201
d: [2,364,1,4,1,0,40,276,206]
s: 364
i: [1,207,0,5,2,0,0,1] t6_boole
s: 207
d: [2,368,1,6,2,1,44,299,207]
s: 368
i: [1,234,0,5,2,0,0,2] t7_boole
s: 234
i: [1,241,0,5,1,0,0,1] rc2_subset_1
s: 241
d: [2,369,1,6,2,1,44,294,207]
s: 369
i: [1,243,0,5,1,0,0,255] rc3_subset_1
s: 243
i: [1,256,0,5,1,0,0,255] fc5_relat_1
s: 256
i: [1,273,0,5,1,0,0,255] rc3_relat_1
s: 273
i: [1,274,0,5,1,0,0,255] rc3_relat_1
s: 274
i: [1,278,0,5,1,0,0,255] rc1_relset_1
s: 278
i: [1,279,0,5,1,0,0,255] rc1_relset_1
s: 279
i: [1,203,0,6,2,0,0,255] rc4_subset_1
s: 203
i: [1,219,0,6,3,0,0,255] cc1_toler_1
s: 219
d: [2,370,0,4,2,0,21,219,206]
s: 370
d: [2,371,1,2,1,1,40,370,324]
a: [3,378,1,2,1,371]
s: 378
i: [1,220,0,6,3,0,0,255] cc1_toler_1
s: 220
d: [2,374,1,2,1,1,40,370,334]
a: [3,383,1,2,1,374]
s: 383
d: [2,386,0,4,2,0,21,220,206]
s: 386
i: [1,221,0,6,3,0,0,255] cc1_toler_1
s: 221
d: [2,387,1,2,1,1,40,386,324]
a: [3,394,1,2,1,387]
s: 394
d: [2,401,0,4,2,0,21,221,206]
s: 401
d: [2,390,1,2,1,1,40,386,334]
a: [3,399,1,2,1,390]
s: 399
i: [1,225,0,6,2,0,0,2] t1_subset
s: 225
i: [1,261,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 261
d: [2,403,1,2,1,1,40,401,324]
a: [3,410,1,2,1,403]
s: 410
i: [1,200,0,7,2,0,0,1] rc1_subset_1
s: 200
d: [2,406,1,2,1,1,40,401,334]
a: [3,415,1,2,1,406]
s: 415
i: [1,202,0,7,2,0,0,255] rc4_subset_1
s: 202
i: [1,232,0,7,2,0,0,255] t3_subset
s: 232
d: [2,372,1,3,1,0,40,370,242]
s: 372
i: [1,233,0,7,3,0,0,1] t8_boole
s: 233
d: [2,388,1,3,1,0,40,386,242]
s: 388
i: [1,245,0,7,1,0,0,255] rc2_partfun1
s: 245
i: [1,257,0,7,1,0,0,255] commutativity_k2_tarski
s: 257
d: [2,404,1,3,1,0,40,401,242]
s: 404
i: [1,205,0,8,3,0,0,2] cc2_relat_1
s: 205
d: [2,421,1,6,2,1,40,205,304]
d: [2,428,1,2,1,1,21,421,255]
a: [3,432,1,2,1,428]
s: 432
i: [1,208,0,8,3,0,0,1] cc1_subset_1
s: 208
i: [1,217,0,8,4,0,0,255] cc3_partfun1
s: 217
d: [2,435,1,6,2,1,40,208,304]
a: [3,445,1,4,1,435]
s: 445
i: [1,226,0,8,3,0,0,2] t2_subset
s: 226
d: [2,456,1,11,2,1,40,226,304]
d: [2,464,1,6,1,1,21,456,197]
a: [3,468,1,6,1,464]
s: 468
i: [1,262,0,8,3,0,0,1] fc10_subset_1
s: 262
i: [1,269,0,8,2,0,0,1] cc1_relset_1
s: 269
d: [2,475,2,6,1,1,40,468,261]
a: [3,481,2,6,1,475]
s: 481
i: [1,275,0,8,1,0,0,255] rc1_relset_1
s: 275
d: [2,373,1,4,1,0,40,370,276]
s: 373
i: [1,280,0,8,1,0,0,255] fc7_relat_1
s: 280
i: [1,209,0,9,3,0,0,255] cc4_subset_1
s: 209
d: [2,389,1,4,1,0,40,386,276]
s: 389
i: [1,210,0,9,4,0,0,255] cc8_relat_1
s: 210
d: [2,405,1,4,1,0,40,401,276]
s: 405
i: [1,213,0,9,4,0,0,255] cc7_relat_1
s: 213
i: [1,236,0,9,3,0,0,2] t5_subset
s: 236
d: [2,426,1,5,2,0,40,205,243]
s: 426
i: [1,270,0,9,2,0,0,255] cc2_relset_1
s: 270
d: [2,528,1,3,1,1,40,270,304]
a: [3,539,1,3,1,528]
s: 539
i: [1,271,0,9,2,0,0,255] cc2_relset_1
s: 271
i: [1,224,0,10,1,0,0,255] d5_tarski
s: 224
d: [2,541,2,6,3,1,40,539,213]
d: [2,542,2,4,2,2,21,541,432]
a: [3,546,2,2,1,542]
s: 546
i: [1,235,0,10,3,0,0,2] t4_subset
s: 235
d: [2,549,1,3,1,1,40,271,304]
a: [3,560,1,3,1,549]
s: 560
i: [1,237,0,10,2,0,0,2] t87_zfmisc_1
s: 237
i: [1,238,0,10,2,0,0,2] t87_zfmisc_1
s: 238
d: [2,571,1,8,2,1,40,235,304]
s: 571
i: [1,258,0,10,3,0,0,1] cc4_relset_1
s: 258
d: [2,580,2,12,3,1,40,571,226]
d: [2,581,2,8,2,2,21,580,445]
s: 581
i: [1,259,0,10,3,0,0,1] cc3_relset_1
s: 259
i: [1,198,0,11,4,0,0,255] cc3_subset_1
s: 198
a: [3,448,1,2,1,435]
s: 448
i: [1,199,0,11,4,0,0,255] cc2_subset_1
s: 199
d: [2,592,3,8,2,2,40,581,261]
s: 592
i: [1,260,0,11,3,0,0,1] cc1_partfun1
s: 260
i: [1,204,0,12,3,0,0,1] t37_relat_1
s: 204
d: [2,627,4,10,2,2,40,592,581]
d: [2,628,4,5,1,2,18,627]
a: [3,632,4,5,1,628]
s: 632
i: [1,267,0,12,4,0,0,255] cc5_relat_1
s: 267
d: [2,651,1,12,3,1,40,267,304]
d: [2,659,1,8,2,1,21,651,255]
s: 659
i: [1,268,0,12,4,0,0,255] cc6_relat_1
s: 268
i: [1,239,0,13,3,0,0,2] t87_zfmisc_1
s: 239
d: [2,663,1,12,3,1,40,268,304]
d: [2,671,1,8,2,1,21,663,255]
s: 671
i: [1,263,0,13,4,0,0,1] cc2_partfun1
s: 263
d: [2,595,3,8,2,2,40,581,238]
s: 595
r: 595
r: 592
r: 581
r: 445
r: 448
i: [1,264,0,19,5,0,0,1] existence_m1_toler_1
s: 264
i: [1,265,0,19,5,0,0,1] rc1_toler_1
s: 265
a: [3,449,1,2,1,435]
s: 449
d: [2,754,2,0,0,2,52,289,449,207]
e: 754
d: [2,755,2,0,0,3,52,324,289,449,233]
e: 755
d: [2,756,2,0,0,3,52,324,289,449,233]
e: 756
r: 449
i: [1,266,0,19,5,0,0,1] rc1_toler_1
s: 266
s: 448
i: [1,229,0,27,7,0,0,1] d2_toler_1
s: 229
i: [1,230,0,27,7,0,0,1] d2_toler_1
s: 230
s: 445
i: [1,231,0,32,8,0,0,1] d2_toler_1
s: 231
d: [2,688,4,9,3,2,40,595,204]
d: [2,689,4,6,2,3,21,688,289]
d: [2,690,4,4,1,4,21,689,432]
a: [3,694,4,4,1,690]
s: 694
i: [1,227,0,33,9,0,0,1] d2_toler_1
s: 227
i: [1,228,0,35,9,0,0,1] d2_toler_1
s: 228
d: [2,814,5,4,1,1,40,694,261]
a: [3,820,5,4,1,814]
s: 820
d: [2,608,1,11,3,1,40,198,304]
d: [2,618,1,7,2,2,21,608,445]
a: [3,622,1,5,1,618]
s: 622
d: [2,815,5,4,1,1,40,694,225]
a: [3,825,5,4,1,815]
s: 825
d: [2,707,1,13,4,1,40,264,304]
d: [2,715,1,10,3,2,21,707,309]
d: [2,716,1,8,2,3,21,715,314]
d: [2,717,1,6,1,4,21,716,319]
a: [3,721,1,6,1,717]
s: 721
d: [2,731,1,13,4,1,40,265,304]
d: [2,739,1,10,3,2,21,731,309]
d: [2,740,1,8,2,3,21,739,314]
d: [2,741,1,6,1,4,21,740,319]
a: [3,745,1,6,1,741]
s: 745
d: [2,765,1,13,4,1,40,266,304]
d: [2,773,1,10,3,2,21,765,309]
d: [2,774,1,8,2,3,21,773,314]
d: [2,775,1,6,1,4,21,774,319]
a: [3,779,1,6,1,775]
s: 779
d: [2,789,1,21,6,1,40,230,304]
d: [2,797,1,18,5,2,21,789,309]
d: [2,798,1,16,4,3,21,797,314]
d: [2,799,1,14,3,4,21,798,319]
s: 799
d: [2,850,2,29,8,4,40,799,229]
d: [2,854,2,21,6,4,18,850]
d: [2,855,2,15,5,4,21,854,304]
d: [2,856,2,12,4,4,21,855,309]
d: [2,857,2,10,3,4,21,856,314]
d: [2,858,2,8,2,4,21,857,319]
a: [3,866,2,4,1,858]
s: 866
d: [2,569,1,4,1,0,44,254,224]
s: 569
d: [2,828,1,27,8,1,40,227,304]
d: [2,836,1,24,7,2,21,828,309]
d: [2,837,1,22,6,3,21,836,314]
d: [2,838,1,20,5,4,21,837,319]
s: 838
d: [2,871,2,16,5,4,40,838,234]
d: [2,872,2,13,4,4,21,871,234]
d: [2,873,2,9,3,5,21,872,368]
d: [2,874,2,5,2,6,21,873,369]
a: [3,877,2,2,1,874]
s: 877
r: 344
r: 329
r: 201
r: 448
r: 546
r: 196
r: 569
r: 254
r: 197
r: 445
r: 262
r: 401
r: 386
r: 370
r: 369
r: 368
r: 263
r: 260
r: 259
r: 258
r: 236
r: 234
r: 233
r: 213
r: 210
r: 209
r: 208
r: 207
r: 206
r: 198
d: [2,886,3,0,0,2,40,877,334]
e: 886
d: [2,885,3,0,0,1,40,877,276]
e: 885
d: [2,884,3,0,0,1,40,877,242]
e: 884
d: [2,883,3,0,0,2,40,877,324]
e: 883
d: [2,882,3,0,0,2,52,324,877]
e: 882
r: 877
s: 329
s: 448
s: 344
s: 196
s: 546
s: 197
a: [3,880,2,3,1,874]
s: 880
r: 838
d: [2,896,5,0,0,2,40,880,694]
e: 896
d: [2,892,5,0,0,2,52,694,880]
e: 892
d: [2,893,6,0,0,3,52,825,546,880,226]
e: 893
f: 892,884,874,756,690,435,428,195,188,189,190,191,192,193,194
% Refutation found. Thanks to Tanya!
% SZS status Theorem for toler_1__t13_toler_1
% SZS output start Proof for toler_1__t13_toler_1
1. ! [X0] : ! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v3_relat_2(X1) & v1_relat_2(X1)) => ((m1_toler_1(k1_xboole_0,X0,X1) & v1_toler_1(k1_xboole_0,X0,X1)) => k1_xboole_0 = X1)) [input t13_toler_1]
2. ~! [X0] : ! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v3_relat_2(X1) & v1_relat_2(X1)) => ((m1_toler_1(k1_xboole_0,X0,X1) & v1_toler_1(k1_xboole_0,X0,X1)) => k1_xboole_0 = X1)) [negated conjecture 1]
7. ! [X0] : (v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_xboole_0(X1))) [input cc1_subset_1]
10. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
25. ! [X0] : ! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v3_relat_2(X1) & v1_relat_2(X1)) => ! [X2] : (m1_toler_1(X2,X0,X1) => (v1_toler_1(X2,X0,X1) <=> ! [X3] : ~(! [X4] : ~(~r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X2)) & r2_hidden(X3,X0) & ~r2_hidden(X3,X2))))) [input d2_toler_1]
39. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
43. ! [X0,X1] : v1_relat_1(k2_zfmisc_1(X0,X1)) [input fc6_relat_1]
52. ! [X0] : ? [X1] : (v1_xboole_0(X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) [input rc2_subset_1]
59. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
60. ! [X0] : (v1_relat_1(X0) => (! [X1] : ! [X2] : ~r2_hidden(k4_tarski(X1,X2),X0) => k1_xboole_0 = X0)) [input t37_relat_1]
62. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
64. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
65. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
66. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [input t87_zfmisc_1]
67. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [input t8_boole]
68. ~! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v3_relat_2(X1) & v1_relat_2(X1)) => ((m1_toler_1(k1_xboole_0,X0,X1) & v1_toler_1(k1_xboole_0,X0,X1)) => k1_xboole_0 = X1)) [flattening 2]
69. ! [X0] : (v1_relat_1(X0) => (! [X1,X2] : ~r2_hidden(k4_tarski(X1,X2),X0) => k1_xboole_0 = X0)) [flattening 60]
72. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 59]
73. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v3_relat_2(X1) & v1_relat_2(X1)) => ! [X2] : (m1_toler_1(X2,X0,X1) => (v1_toler_1(X2,X0,X1) <=> ! [X3] : ~(! [X4] : ~(~r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X2)) & r2_hidden(X3,X0) & ~r2_hidden(X3,X2))))) [flattening 25]
75. ! [X0,X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [flattening 67]
76. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 65]
77. ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 62]
79. ! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [flattening 66]
93. ? [X0,X1] : ((k1_xboole_0 != X1 & (m1_toler_1(k1_xboole_0,X0,X1) & v1_toler_1(k1_xboole_0,X0,X1))) & (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v3_relat_2(X1) & v1_relat_2(X1))) [ennf transformation 68]
94. ? [X0,X1] : (k1_xboole_0 != X1 & m1_toler_1(k1_xboole_0,X0,X1) & v1_toler_1(k1_xboole_0,X0,X1) & m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v3_relat_2(X1) & v1_relat_2(X1)) [flattening 93]
101. ! [X0] : ((k1_xboole_0 = X0 | ? [X1,X2] : r2_hidden(k4_tarski(X1,X2),X0)) | ~v1_relat_1(X0)) [ennf transformation 69]
102. ! [X0] : (k1_xboole_0 = X0 | ? [X1,X2] : r2_hidden(k4_tarski(X1,X2),X0) | ~v1_relat_1(X0)) [flattening 101]
103. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 10]
105. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 64]
106. ! [X0] : (! [X1] : (v1_xboole_0(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_xboole_0(X0)) [ennf transformation 7]
121. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 72]
122. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 121]
123. ! [X0,X1] : (! [X2] : ((v1_toler_1(X2,X0,X1) <=> ! [X3] : (? [X4] : (~r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X2)) | ~r2_hidden(X3,X0) | r2_hidden(X3,X2))) | ~m1_toler_1(X2,X0,X1)) | (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~v1_partfun1(X1,X0) | ~v3_relat_2(X1) | ~v1_relat_2(X1))) [ennf transformation 73]
124. ! [X0,X1] : (! [X2] : ((v1_toler_1(X2,X0,X1) <=> ! [X3] : (? [X4] : (~r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X2)) | ~r2_hidden(X3,X0) | r2_hidden(X3,X2))) | ~m1_toler_1(X2,X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~v1_partfun1(X1,X0) | ~v3_relat_2(X1) | ~v1_relat_2(X1)) [flattening 123]
126. ! [X0,X1] : (~v1_xboole_0(X1) | X0 = X1 | ~v1_xboole_0(X0)) [ennf transformation 75]
127. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 76]
128. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))) [ennf transformation 77]
129. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [flattening 128]
149. ? [X0,X1] : (k1_xboole_0 != X1 & m1_toler_1(k1_xboole_0,X0,X1) & v1_toler_1(k1_xboole_0,X0,X1) & m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) & v1_partfun1(X1,X0) & v3_relat_2(X1) & v1_relat_2(X1)) => (k1_xboole_0 != sK1 & m1_toler_1(k1_xboole_0,sK0,sK1) & v1_toler_1(k1_xboole_0,sK0,sK1) & m1_subset_1(sK1,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v1_partfun1(sK1,sK0) & v3_relat_2(sK1) & v1_relat_2(sK1)) [choice axiom]
150. k1_xboole_0 != sK1 & m1_toler_1(k1_xboole_0,sK0,sK1) & v1_toler_1(k1_xboole_0,sK0,sK1) & m1_subset_1(sK1,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) & v1_partfun1(sK1,sK0) & v3_relat_2(sK1) & v1_relat_2(sK1) [skolemisation 94,149]
155. ! [X0] : (? [X1,X2] : r2_hidden(k4_tarski(X1,X2),X0) => r2_hidden(k4_tarski(sK4(X0),sK5(X0)),X0)) [choice axiom]
156. ! [X0] : (k1_xboole_0 = X0 | r2_hidden(k4_tarski(sK4(X0),sK5(X0)),X0) | ~v1_relat_1(X0)) [skolemisation 102,155]
157. ! [X0,X1] : (! [X2] : (((v1_toler_1(X2,X0,X1) | ? [X3] : (! [X4] : (r2_hidden(k4_tarski(X3,X4),X1) | ~r2_hidden(X4,X2)) & r2_hidden(X3,X0) & ~r2_hidden(X3,X2))) & (! [X3] : (? [X4] : (~r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X2)) | ~r2_hidden(X3,X0) | r2_hidden(X3,X2)) | ~v1_toler_1(X2,X0,X1))) | ~m1_toler_1(X2,X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~v1_partfun1(X1,X0) | ~v3_relat_2(X1) | ~v1_relat_2(X1)) [nnf transformation 124]
158. ! [X0,X1] : (! [X2] : (((v1_toler_1(X2,X0,X1) | ? [X3] : (! [X4] : (r2_hidden(k4_tarski(X3,X4),X1) | ~r2_hidden(X4,X2)) & r2_hidden(X3,X0) & ~r2_hidden(X3,X2))) & (! [X5] : (? [X6] : (~r2_hidden(k4_tarski(X5,X6),X1) & r2_hidden(X6,X2)) | ~r2_hidden(X5,X0) | r2_hidden(X5,X2)) | ~v1_toler_1(X2,X0,X1))) | ~m1_toler_1(X2,X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~v1_partfun1(X1,X0) | ~v3_relat_2(X1) | ~v1_relat_2(X1)) [rectify 157]
159. ! [X2,X1,X0] : (? [X3] : (! [X4] : (r2_hidden(k4_tarski(X3,X4),X1) | ~r2_hidden(X4,X2)) & r2_hidden(X3,X0) & ~r2_hidden(X3,X2)) => (! [X4] : (r2_hidden(k4_tarski(sK6(X0,X1,X2),X4),X1) | ~r2_hidden(X4,X2)) & r2_hidden(sK6(X0,X1,X2),X0) & ~r2_hidden(sK6(X0,X1,X2),X2))) [choice axiom]
160. ! [X5,X2,X1] : (? [X6] : (~r2_hidden(k4_tarski(X5,X6),X1) & r2_hidden(X6,X2)) => (~r2_hidden(k4_tarski(X5,sK7(X1,X2,X5)),X1) & r2_hidden(sK7(X1,X2,X5),X2))) [choice axiom]
161. ! [X0,X1] : (! [X2] : (((v1_toler_1(X2,X0,X1) | (! [X4] : (r2_hidden(k4_tarski(sK6(X0,X1,X2),X4),X1) | ~r2_hidden(X4,X2)) & r2_hidden(sK6(X0,X1,X2),X0) & ~r2_hidden(sK6(X0,X1,X2),X2))) & (! [X5] : ((~r2_hidden(k4_tarski(X5,sK7(X1,X2,X5)),X1) & r2_hidden(sK7(X1,X2,X5),X2)) | ~r2_hidden(X5,X0) | r2_hidden(X5,X2)) | ~v1_toler_1(X2,X0,X1))) | ~m1_toler_1(X2,X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~v1_partfun1(X1,X0) | ~v3_relat_2(X1) | ~v1_relat_2(X1)) [skolemisation 158,160,159]
162. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | (~r2_hidden(X1,X3) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [nnf transformation 79]
163. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [flattening 162]
166. ! [X0] : (? [X1] : (v1_xboole_0(X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (v1_xboole_0(sK9(X0)) & m1_subset_1(sK9(X0),k1_zfmisc_1(X0)))) [choice axiom]
167. ! [X0] : (v1_xboole_0(sK9(X0)) & m1_subset_1(sK9(X0),k1_zfmisc_1(X0))) [skolemisation 52,166]
188. v1_relat_2(sK1) [cnf transformation 150] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
189. v3_relat_2(sK1) [cnf transformation 150] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
190. v1_partfun1(sK1,sK0) [cnf transformation 150] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
191. m1_subset_1(sK1,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) [cnf transformation 150] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
192. v1_toler_1(k1_xboole_0,sK0,sK1) [cnf transformation 150] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
193. m1_toler_1(k1_xboole_0,sK0,sK1) [cnf transformation 150] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
194. k1_xboole_0 != sK1 [cnf transformation 150] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
195. v1_xboole_0(k1_xboole_0) [cnf transformation 39] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
204. r2_hidden(k4_tarski(sK4(X0),sK5(X0)),X0) | k1_xboole_0 = X0 | ~v1_relat_1(X0) [cnf transformation 156] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
205. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 103] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
207. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 105] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
208. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 106] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
226. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 122] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
227. ~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X0,X0))) | ~r2_hidden(X5,X0) | r2_hidden(X5,X2) | ~v1_toler_1(X2,X0,X1) | ~m1_toler_1(X2,X0,X1) | r2_hidden(sK7(X1,X2,X5),X2) | ~v1_partfun1(X1,X0) | ~v3_relat_2(X1) | ~v1_relat_2(X1) [cnf transformation 161] {a:0,w:33,wCS:13200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
233. X0 = X1 | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 126] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
234. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 127] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
235. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1) [cnf transformation 129] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
238. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X1,X3) [cnf transformation 163] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
242. v1_xboole_0(sK9(X0)) [cnf transformation 167] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
255. v1_relat_1(k2_zfmisc_1(X0,X1)) [cnf transformation 43] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
287. 1 <=> k1_xboole_0 = sK1 [avatar definition]
289. k1_xboole_0 != sK1 <- (~1) [avatar component clause 287] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
290. ~1 [avatar split clause 194,287]
292. 2 <=> m1_toler_1(k1_xboole_0,sK0,sK1) [avatar definition]
294. m1_toler_1(k1_xboole_0,sK0,sK1) <- (2) [avatar component clause 292] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
295. 2 [avatar split clause 193,292]
297. 3 <=> v1_toler_1(k1_xboole_0,sK0,sK1) [avatar definition]
299. v1_toler_1(k1_xboole_0,sK0,sK1) <- (3) [avatar component clause 297] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
300. 3 [avatar split clause 192,297]
302. 4 <=> m1_subset_1(sK1,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) [avatar definition]
304. m1_subset_1(sK1,k1_zfmisc_1(k2_zfmisc_1(sK0,sK0))) <- (4) [avatar component clause 302] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
305. 4 [avatar split clause 191,302]
307. 5 <=> v1_partfun1(sK1,sK0) [avatar definition]
309. v1_partfun1(sK1,sK0) <- (5) [avatar component clause 307] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
310. 5 [avatar split clause 190,307]
312. 6 <=> v3_relat_2(sK1) [avatar definition]
314. v3_relat_2(sK1) <- (6) [avatar component clause 312] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
315. 6 [avatar split clause 189,312]
317. 7 <=> v1_relat_2(sK1) [avatar definition]
319. v1_relat_2(sK1) <- (7) [avatar component clause 317] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
320. 7 [avatar split clause 188,317]
322. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
324. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 322] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
325. 8 [avatar split clause 195,322]
368. v1_toler_1(X0,sK0,sK1) | ~v1_xboole_0(X0) <- (3) [superposition 299,207] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
369. m1_toler_1(X0,sK0,sK1) | ~v1_xboole_0(X0) <- (2) [superposition 294,207] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
421. v1_relat_1(sK1) | ~v1_relat_1(k2_zfmisc_1(sK0,sK0)) <- (4) [resolution 205,304] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
428. v1_relat_1(sK1) <- (4) [subsumption resolution 421,255] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
430. 22 <=> v1_relat_1(sK1) [avatar definition]
432. v1_relat_1(sK1) <- (22) [avatar component clause 430] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
433. 22 | ~4 [avatar split clause 428,302,430]
435. v1_xboole_0(sK1) | ~v1_xboole_0(k2_zfmisc_1(sK0,sK0)) <- (4) [resolution 208,304] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
443. 23 <=> v1_xboole_0(k2_zfmisc_1(sK0,sK0)) [avatar definition]
445. ~v1_xboole_0(k2_zfmisc_1(sK0,sK0)) <- (~23) [avatar component clause 443] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
447. 24 <=> v1_xboole_0(sK1) [avatar definition]
449. v1_xboole_0(sK1) <- (24) [avatar component clause 447] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
450. ~23 | 24 | ~4 [avatar split clause 435,302,447,443]
571. m1_subset_1(X3,k2_zfmisc_1(sK0,sK0)) | ~r2_hidden(X3,sK1) <- (4) [resolution 235,304] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
580. ~r2_hidden(X0,sK1) | v1_xboole_0(k2_zfmisc_1(sK0,sK0)) | r2_hidden(X0,k2_zfmisc_1(sK0,sK0)) <- (4) [resolution 571,226] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
581. r2_hidden(X0,k2_zfmisc_1(sK0,sK0)) | ~r2_hidden(X0,sK1) <- (4, ~23) [subsumption resolution 580,445] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
595. ~r2_hidden(k4_tarski(X3,X4),sK1) | r2_hidden(X4,sK0) <- (4, ~23) [resolution 581,238] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
688. r2_hidden(sK5(sK1),sK0) | k1_xboole_0 = sK1 | ~v1_relat_1(sK1) <- (4, ~23) [resolution 595,204] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
689. r2_hidden(sK5(sK1),sK0) | ~v1_relat_1(sK1) <- (~1, 4, ~23) [subsumption resolution 688,289] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
690. r2_hidden(sK5(sK1),sK0) <- (~1, 4, 22, ~23) [subsumption resolution 689,432] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
692. 32 <=> r2_hidden(sK5(sK1),sK0) [avatar definition]
694. r2_hidden(sK5(sK1),sK0) <- (32) [avatar component clause 692] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
695. 32 | 1 | ~4 | ~22 | 23 [avatar split clause 690,443,430,302,287,692]
756. $false <- (~1, 8, 24) [unit resulting resolution 324,289,449,233] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
763. 1 | ~8 | ~24 [avatar contradiction clause 756]
828. ~r2_hidden(X4,sK0) | r2_hidden(X4,X5) | ~v1_toler_1(X5,sK0,sK1) | ~m1_toler_1(X5,sK0,sK1) | r2_hidden(sK7(sK1,X5,X4),X5) | ~v1_partfun1(sK1,sK0) | ~v3_relat_2(sK1) | ~v1_relat_2(sK1) <- (4) [resolution 227,304] {a:1,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
836. ~r2_hidden(X4,sK0) | r2_hidden(X4,X5) | ~v1_toler_1(X5,sK0,sK1) | ~m1_toler_1(X5,sK0,sK1) | r2_hidden(sK7(sK1,X5,X4),X5) | ~v3_relat_2(sK1) | ~v1_relat_2(sK1) <- (4, 5) [subsumption resolution 828,309] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
837. ~r2_hidden(X4,sK0) | r2_hidden(X4,X5) | ~v1_toler_1(X5,sK0,sK1) | ~m1_toler_1(X5,sK0,sK1) | r2_hidden(sK7(sK1,X5,X4),X5) | ~v1_relat_2(sK1) <- (4, 5, 6) [subsumption resolution 836,314] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
838. r2_hidden(sK7(sK1,X5,X4),X5) | r2_hidden(X4,X5) | ~v1_toler_1(X5,sK0,sK1) | ~m1_toler_1(X5,sK0,sK1) | ~r2_hidden(X4,sK0) <- (4, 5, 6, 7) [subsumption resolution 837,319] {a:1,w:20,wCS:2000,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
871. r2_hidden(X4,X5) | ~v1_toler_1(X5,sK0,sK1) | ~m1_toler_1(X5,sK0,sK1) | ~r2_hidden(X4,sK0) | ~v1_xboole_0(X5) <- (4, 5, 6, 7) [resolution 838,234] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
872. ~v1_toler_1(X5,sK0,sK1) | ~m1_toler_1(X5,sK0,sK1) | ~r2_hidden(X4,sK0) | ~v1_xboole_0(X5) <- (4, 5, 6, 7) [subsumption resolution 871,234] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
873. ~m1_toler_1(X5,sK0,sK1) | ~r2_hidden(X4,sK0) | ~v1_xboole_0(X5) <- (3, 4, 5, 6, 7) [subsumption resolution 872,368] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
874. ~r2_hidden(X4,sK0) | ~v1_xboole_0(X5) <- (2, 3, 4, 5, 6, 7) [subsumption resolution 873,369] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
876. 42 <=> ! [X5] : ~v1_xboole_0(X5) [avatar definition]
877. ~v1_xboole_0(X5) <- (42) [avatar component clause 876] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
879. 43 <=> ! [X4] : ~r2_hidden(X4,sK0) [avatar definition]
880. ~r2_hidden(X4,sK0) <- (43) [avatar component clause 879] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
881. 42 | 43 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 874,317,312,307,302,297,292,879,876]
884. $false <- (42) [resolution 877,242] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
889. ~42 [avatar contradiction clause 884]
892. $false <- (32, 43) [unit resulting resolution 694,880] {a:5,w:0,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
899. ~32 | ~43 [avatar contradiction clause 892]
901. $false [avatar sat refutation 290,295,300,305,310,315,320,325,433,450,695,763,881,889,899] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:349,thDist:-349}
% SZS output end Proof for toler_1__t13_toler_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 66
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 7
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 98
% Generated clauses: 571
% Active clauses: 144
% Passive clauses: 346
% Final active clauses: 106
% Final passive clauses: 195
% Main loop iterations started: 149
% 
% >>> Simplifying Inferences
% Duplicate literals: 7
% Fw subsumption resolutions: 93
% 
% >>> Deletion Inferences
% Simple tautologies: 19
% Forward subsumptions: 56
% Backward subsumptions: 31
% 
% >>> Generating Inferences
% Binary resolution: 270
% Unit resulting resolution: 6
% Forward superposition: 9
% Backward superposition: 9
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 19
% Sat splitting refutations: 11
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 55
% SAT solver unit clauses: 15
% SAT solver binary clauses: 21
% 
% Memory used [KB]: 5500
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% binary resolution index maintenance: 0.001 s
% splitting component index usage: 0.001 s
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.007 s ( own 0.001 s ) 

