i: [1,167,0,5,2,1,0,0] l3_wellord1
a: [3,246,0,2,1,167]
s: 246
i: [1,163,0,2,1,1,0,0] l3_wellord1
a: [3,269,0,2,1,163]
s: 269
i: [1,168,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,274,0,2,1,168]
s: 274
i: [1,227,0,2,1,0,0,255] rc7_funct_1
a: [3,284,0,2,1,227]
s: 284
i: [1,226,0,2,1,0,0,255] rc7_funct_1
a: [3,289,0,2,1,226]
s: 289
i: [1,229,0,2,1,0,0,2] rc1_relat_1
a: [3,294,0,2,1,229]
s: 294
i: [1,228,0,2,1,0,0,2] rc1_relat_1
a: [3,299,0,2,1,228]
s: 299
i: [1,232,0,2,1,0,0,255] rc4_funct_1
a: [3,304,0,2,1,232]
s: 304
i: [1,231,0,2,1,0,0,255] rc4_funct_1
a: [3,309,0,2,1,231]
s: 309
i: [1,230,0,2,1,0,0,255] rc4_funct_1
a: [3,314,0,2,1,230]
s: 314
i: [1,233,0,2,1,0,0,255] rc2_relat_1
a: [3,319,0,2,1,233]
s: 319
i: [1,235,0,2,1,0,0,255] rc3_funct_1
a: [3,324,0,2,1,235]
s: 324
i: [1,234,0,2,1,0,0,255] rc3_funct_1
a: [3,329,0,2,1,234]
s: 329
i: [1,237,0,2,1,0,0,4] rc1_funct_1
a: [3,334,0,2,1,237]
s: 334
i: [1,236,0,2,1,0,0,4] rc1_funct_1
a: [3,339,0,2,1,236]
s: 339
i: [1,240,0,2,1,0,0,255] rc5_funct_1
a: [3,344,0,2,1,240]
s: 344
i: [1,239,0,2,1,0,0,255] rc5_funct_1
a: [3,349,0,2,1,239]
s: 349
i: [1,238,0,2,1,0,0,255] rc5_funct_1
a: [3,354,0,2,1,238]
s: 354
i: [1,242,0,2,1,0,0,255] rc2_funct_1
a: [3,359,0,2,1,242]
s: 359
i: [1,241,0,2,1,0,0,255] rc2_funct_1
a: [3,364,0,2,1,241]
s: 364
i: [1,169,0,3,1,0,0,255] d2_xboole_0
a: [3,279,0,3,1,169]
s: 279
r: 274
d: [2,366,0,2,1,2,23,274,279]
a: [3,370,0,2,1,366]
s: 370
i: [1,164,0,15,4,1,0,0] l3_wellord1
a: [3,264,0,13,3,164]
s: 264
i: [1,173,0,4,2,0,0,255] cc7_funct_1
s: 173
i: [1,174,0,4,2,0,0,4] cc1_funct_1
s: 174
i: [1,175,0,4,2,0,0,2] cc1_relat_1
s: 175
d: [2,372,1,2,1,1,40,173,370]
a: [3,376,1,2,1,372]
s: 376
i: [1,216,0,4,1,0,0,255] existence_m1_subset_1
s: 216
d: [2,378,1,2,1,1,40,174,370]
a: [3,382,1,2,1,378]
s: 382
i: [1,170,0,5,1,0,0,255] t1_boole
s: 170
i: [1,176,0,5,2,0,0,255] t6_boole
s: 176
d: [2,384,1,2,1,1,40,175,370]
a: [3,388,1,2,1,384]
s: 388
i: [1,177,0,5,2,0,0,3] fc10_relat_1
s: 177
i: [1,178,0,5,2,0,0,3] fc11_relat_1
s: 178
i: [1,213,0,5,2,0,0,1] t7_boole
s: 213
i: [1,217,0,5,1,0,0,3] idempotence_k2_xboole_0
s: 217
i: [1,218,0,5,1,0,0,255] fc1_funct_1
s: 218
i: [1,219,0,5,1,0,0,255] fc5_relat_1
s: 219
d: [2,396,1,5,2,0,40,177,175]
s: 396
i: [1,211,0,6,2,0,0,255] t1_subset
s: 211
i: [1,222,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 222
d: [2,397,1,5,2,0,40,177,174]
s: 397
i: [1,171,0,7,3,0,0,255] cc8_funct_1
s: 171
d: [2,398,1,5,2,0,40,177,173]
s: 398
i: [1,172,0,7,3,0,0,255] cc8_funct_1
s: 172
i: [1,189,0,7,3,0,0,3] fc8_relat_1
s: 189
d: [2,399,1,5,2,0,40,178,175]
s: 399
i: [1,190,0,7,3,0,0,3] fc9_relat_1
s: 190
d: [2,400,1,5,2,0,40,178,174]
s: 400
i: [1,194,0,7,3,0,0,3] fc24_relat_1
s: 194
i: [1,195,0,7,3,0,0,3] fc25_relat_1
s: 195
d: [2,401,1,5,2,0,40,178,173]
s: 401
i: [1,208,0,7,3,0,0,255] fc14_funct_1
s: 208
d: [2,402,1,5,2,0,40,171,216]
s: 402
i: [1,212,0,7,3,0,0,1] t8_boole
s: 212
i: [1,220,0,7,1,0,0,255] commutativity_k2_tarski
s: 220
d: [2,403,1,5,2,0,40,172,216]
s: 403
i: [1,221,0,7,1,0,0,3] commutativity_k2_xboole_0
s: 221
d: [2,406,1,5,1,0,44,221,170]
s: 406
i: [1,180,0,8,3,0,0,1] d12_relat_2
s: 180
i: [1,181,0,8,3,0,0,1] d12_relat_2
s: 181
d: [2,407,1,5,1,0,44,221,170]
s: 407
i: [1,193,0,8,4,0,0,255] cc6_funct_1
s: 193
d: [2,408,1,5,1,0,44,170,221]
s: 408
i: [1,200,0,8,4,0,0,255] cc4_funct_1
s: 200
d: [2,419,0,6,3,0,21,200,174]
d: [2,420,0,4,2,0,21,419,175]
s: 420
d: [2,421,1,2,1,1,40,420,370]
a: [3,427,1,2,1,421]
s: 427
i: [1,204,0,8,4,0,0,255] cc5_funct_1
s: 204
d: [2,409,1,5,1,0,44,170,221]
s: 409
i: [1,210,0,8,3,0,0,255] t2_subset
s: 210
i: [1,224,0,8,3,0,0,3] fc3_relat_1
s: 224
d: [2,422,1,5,2,0,40,420,177]
s: 422
i: [1,225,0,8,1,0,0,255] fc7_relat_1
s: 225
d: [2,423,1,5,2,0,40,420,178]
s: 423
i: [1,188,0,9,4,0,0,3] fc18_funct_1
s: 188
i: [1,203,0,9,4,0,0,255] fc13_funct_1
s: 203
d: [2,430,1,6,2,0,40,210,216]
s: 430
i: [1,207,0,9,4,0,0,255] fc12_funct_1
s: 207
d: [2,443,2,6,2,0,40,430,222]
s: 443
i: [1,179,0,10,2,0,0,2] d6_relat_1
s: 179
i: [1,183,0,10,3,0,0,2] d4_relat_2
s: 183
d: [2,390,1,7,3,0,44,176,176]
s: 390
i: [1,184,0,10,3,0,0,2] d4_relat_2
s: 184
d: [2,391,1,7,2,0,44,170,176]
s: 391
i: [1,209,0,10,1,0,0,255] d5_tarski
s: 209
i: [1,214,0,11,3,0,0,2] t15_relat_1
s: 214
d: [2,452,1,7,3,0,40,183,213]
s: 452
i: [1,215,0,11,3,0,0,2] t15_relat_1
s: 215
d: [2,455,1,7,3,0,40,184,213]
s: 455
i: [1,187,0,12,3,0,0,2] d4_relat_2
s: 187
i: [1,223,0,12,5,0,0,255] fc15_funct_1
s: 223
d: [2,410,2,7,2,0,44,406,176]
s: 410
i: [1,185,0,14,3,0,0,2] d4_relat_2
s: 185
d: [2,488,1,21,4,1,40,185,264]
d: [2,496,1,19,3,2,21,488,269]
s: 496
i: [1,186,0,14,3,0,0,2] d4_relat_2
s: 186
i: [1,182,0,24,7,0,0,2] d4_relat_2
s: 182
d: [2,498,2,15,4,2,40,186,496]
d: [2,507,2,12,3,2,18,498]
d: [2,508,2,10,2,2,21,507,269]
s: 508
r: 496
d: [2,499,1,21,4,1,40,186,264]
d: [2,509,1,19,3,2,21,499,269]
s: 509
d: [2,521,3,15,4,2,44,187,508]
d: [2,524,3,12,3,2,18,521]
d: [2,525,3,5,2,2,19,524]
d: [2,530,3,3,1,2,21,525,269]
s: 530
r: 508
d: [2,531,4,0,0,3,52,269,246,530,181]
e: 531
d: [2,532,4,4,2,2,40,530,181]
d: [2,534,4,2,1,3,21,532,246]
d: [2,535,4,0,0,3,21,534,269]
e: 535
r: 246
a: [3,245,0,2,1,167]
s: 245
a: [3,250,0,3,1,167]
s: 250
i: [1,166,0,7,2,1,0,0] l3_wellord1
a: [3,255,0,5,1,166]
s: 255
r: 530
r: 264
i: [1,165,0,7,2,1,0,0] l3_wellord1
a: [3,260,0,5,1,165]
s: 260
d: [2,545,1,2,1,1,40,255,213]
a: [3,580,1,2,1,545]
s: 580
d: [2,541,1,6,2,1,40,255,215]
d: [2,555,1,4,1,2,21,541,269]
a: [3,559,1,4,1,555]
s: 559
d: [2,542,1,6,2,1,40,255,214]
d: [2,561,1,4,1,2,21,542,269]
a: [3,565,1,4,1,561]
s: 565
d: [2,614,2,3,1,1,40,559,213]
a: [3,628,2,3,1,614]
s: 628
d: [2,543,1,5,1,1,40,255,222]
a: [3,570,1,5,1,543]
s: 570
d: [2,612,2,4,1,1,40,559,222]
a: [3,618,2,4,1,612]
s: 618
d: [2,544,1,5,1,1,40,255,211]
a: [3,575,1,5,1,544]
s: 575
d: [2,590,1,5,1,1,40,260,222]
a: [3,604,1,5,1,590]
s: 604
d: [2,645,2,6,2,1,40,575,172]
a: [3,650,2,2,1,645]
s: 650
d: [2,591,1,5,1,1,40,260,211]
a: [3,609,1,5,1,591]
s: 609
d: [2,613,2,4,1,1,40,559,211]
a: [3,623,2,4,1,613]
s: 623
d: [2,540,1,19,6,1,40,255,182]
d: [2,548,1,16,5,2,21,540,250]
d: [2,582,1,14,4,3,21,548,269]
a: [3,585,1,9,3,582]
s: 585
d: [2,681,2,12,4,1,40,585,180]
d: [2,683,2,8,3,2,21,681,565]
d: [2,684,2,4,2,3,21,683,559]
d: [2,685,2,2,1,4,21,684,245]
d: [2,686,2,0,0,5,21,685,269]
e: 686
d: [2,537,4,2,1,2,21,532,269]
d: [2,587,1,19,6,1,40,260,182]
d: [2,593,1,16,5,2,21,587,250]
d: [2,594,1,11,4,3,21,593,255]
d: [2,595,1,9,3,4,21,594,269]
d: [2,588,1,6,2,1,40,260,215]
d: [2,597,1,4,1,2,21,588,269]
d: [2,589,1,6,2,1,40,260,214]
d: [2,599,1,4,1,2,21,589,269]
f: 686,599,597,595,537,163,164,165,166,167
% Refutation found. Thanks to Tanya!
% SZS status Theorem for wellord1__l3_wellord1
% SZS output start Proof for wellord1__l3_wellord1
1. ! [X0] : (v1_relat_1(X0) => (v4_relat_2(X0) <=> ! [X1] : ! [X2] : ((r2_hidden(k4_tarski(X2,X1),X0) & r2_hidden(k4_tarski(X1,X2),X0)) => X1 = X2))) [input l3_wellord1]
2. ~! [X0] : (v1_relat_1(X0) => (v4_relat_2(X0) <=> ! [X1] : ! [X2] : ((r2_hidden(k4_tarski(X2,X1),X0) & r2_hidden(k4_tarski(X1,X2),X0)) => X1 = X2))) [negated conjecture 1]
16. ! [X0] : (v1_relat_1(X0) => (v4_relat_2(X0) <=> r4_relat_2(X0,k1_relat_1(X0)))) [input d12_relat_2]
18. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r4_relat_2(X0,X1) <=> ! [X2] : ! [X3] : ((r2_hidden(k4_tarski(X3,X2),X0) & r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(X3,X1) & r2_hidden(X2,X1)) => X2 = X3))) [input d4_relat_2]
57. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r2_hidden(k4_tarski(X0,X1),X2) => (r2_hidden(X1,k1_relat_1(X2)) & r2_hidden(X0,k1_relat_1(X2))))) [input t15_relat_1]
64. ~! [X0] : (v1_relat_1(X0) => (v4_relat_2(X0) <=> ! [X1,X2] : ((r2_hidden(k4_tarski(X2,X1),X0) & r2_hidden(k4_tarski(X1,X2),X0)) => X1 = X2))) [flattening 2]
65. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r4_relat_2(X0,X1) <=> ! [X2,X3] : ((r2_hidden(k4_tarski(X3,X2),X0) & r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(X3,X1) & r2_hidden(X2,X1)) => X2 = X3))) [flattening 18]
71. ! [X0,X1,X2] : (v1_relat_1(X2) => (r2_hidden(k4_tarski(X0,X1),X2) => (r2_hidden(X1,k1_relat_1(X2)) & r2_hidden(X0,k1_relat_1(X2))))) [flattening 57]
81. ? [X0] : ((v4_relat_2(X0) <~> ! [X1,X2] : (X1 = X2 | (~r2_hidden(k4_tarski(X2,X1),X0) | ~r2_hidden(k4_tarski(X1,X2),X0)))) & v1_relat_1(X0)) [ennf transformation 64]
82. ? [X0] : ((v4_relat_2(X0) <~> ! [X1,X2] : (X1 = X2 | ~r2_hidden(k4_tarski(X2,X1),X0) | ~r2_hidden(k4_tarski(X1,X2),X0))) & v1_relat_1(X0)) [flattening 81]
91. ! [X0] : ((v4_relat_2(X0) <=> r4_relat_2(X0,k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 16]
92. ! [X0] : (! [X1] : (r4_relat_2(X0,X1) <=> ! [X2,X3] : (X2 = X3 | (~r2_hidden(k4_tarski(X3,X2),X0) | ~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(X3,X1) | ~r2_hidden(X2,X1)))) | ~v1_relat_1(X0)) [ennf transformation 65]
93. ! [X0] : (! [X1] : (r4_relat_2(X0,X1) <=> ! [X2,X3] : (X2 = X3 | ~r2_hidden(k4_tarski(X3,X2),X0) | ~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(X3,X1) | ~r2_hidden(X2,X1))) | ~v1_relat_1(X0)) [flattening 92]
127. ! [X0,X1,X2] : (((r2_hidden(X1,k1_relat_1(X2)) & r2_hidden(X0,k1_relat_1(X2))) | ~r2_hidden(k4_tarski(X0,X1),X2)) | ~v1_relat_1(X2)) [ennf transformation 71]
128. ! [X0,X1,X2] : ((r2_hidden(X1,k1_relat_1(X2)) & r2_hidden(X0,k1_relat_1(X2))) | ~r2_hidden(k4_tarski(X0,X1),X2) | ~v1_relat_1(X2)) [flattening 127]
134. ? [X0] : (((? [X1,X2] : (X1 != X2 & r2_hidden(k4_tarski(X2,X1),X0) & r2_hidden(k4_tarski(X1,X2),X0)) | ~v4_relat_2(X0)) & (! [X1,X2] : (X1 = X2 | ~r2_hidden(k4_tarski(X2,X1),X0) | ~r2_hidden(k4_tarski(X1,X2),X0)) | v4_relat_2(X0))) & v1_relat_1(X0)) [nnf transformation 82]
135. ? [X0] : ((? [X1,X2] : (X1 != X2 & r2_hidden(k4_tarski(X2,X1),X0) & r2_hidden(k4_tarski(X1,X2),X0)) | ~v4_relat_2(X0)) & (! [X1,X2] : (X1 = X2 | ~r2_hidden(k4_tarski(X2,X1),X0) | ~r2_hidden(k4_tarski(X1,X2),X0)) | v4_relat_2(X0)) & v1_relat_1(X0)) [flattening 134]
136. ? [X0] : ((? [X1,X2] : (X1 != X2 & r2_hidden(k4_tarski(X2,X1),X0) & r2_hidden(k4_tarski(X1,X2),X0)) | ~v4_relat_2(X0)) & (! [X3,X4] : (X3 = X4 | ~r2_hidden(k4_tarski(X4,X3),X0) | ~r2_hidden(k4_tarski(X3,X4),X0)) | v4_relat_2(X0)) & v1_relat_1(X0)) [rectify 135]
137. ? [X0] : ((? [X1,X2] : (X1 != X2 & r2_hidden(k4_tarski(X2,X1),X0) & r2_hidden(k4_tarski(X1,X2),X0)) | ~v4_relat_2(X0)) & (! [X3,X4] : (X3 = X4 | ~r2_hidden(k4_tarski(X4,X3),X0) | ~r2_hidden(k4_tarski(X3,X4),X0)) | v4_relat_2(X0)) & v1_relat_1(X0)) => ((? [X2,X1] : (X1 != X2 & r2_hidden(k4_tarski(X2,X1),sK0) & r2_hidden(k4_tarski(X1,X2),sK0)) | ~v4_relat_2(sK0)) & (! [X4,X3] : (X3 = X4 | ~r2_hidden(k4_tarski(X4,X3),sK0) | ~r2_hidden(k4_tarski(X3,X4),sK0)) | v4_relat_2(sK0)) & v1_relat_1(sK0)) [choice axiom]
138. ? [X2,X1] : (X1 != X2 & r2_hidden(k4_tarski(X2,X1),sK0) & r2_hidden(k4_tarski(X1,X2),sK0)) => (sK1 != sK2 & r2_hidden(k4_tarski(sK2,sK1),sK0) & r2_hidden(k4_tarski(sK1,sK2),sK0)) [choice axiom]
139. ((sK1 != sK2 & r2_hidden(k4_tarski(sK2,sK1),sK0) & r2_hidden(k4_tarski(sK1,sK2),sK0)) | ~v4_relat_2(sK0)) & (! [X3,X4] : (X3 = X4 | ~r2_hidden(k4_tarski(X4,X3),sK0) | ~r2_hidden(k4_tarski(X3,X4),sK0)) | v4_relat_2(sK0)) & v1_relat_1(sK0) [skolemisation 136,138,137]
140. ! [X0] : (((v4_relat_2(X0) | ~r4_relat_2(X0,k1_relat_1(X0))) & (r4_relat_2(X0,k1_relat_1(X0)) | ~v4_relat_2(X0))) | ~v1_relat_1(X0)) [nnf transformation 91]
141. ! [X0] : (! [X1] : ((r4_relat_2(X0,X1) | ? [X2,X3] : (X2 != X3 & r2_hidden(k4_tarski(X3,X2),X0) & r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(X3,X1) & r2_hidden(X2,X1))) & (! [X2,X3] : (X2 = X3 | ~r2_hidden(k4_tarski(X3,X2),X0) | ~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(X3,X1) | ~r2_hidden(X2,X1)) | ~r4_relat_2(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 93]
142. ! [X0] : (! [X1] : ((r4_relat_2(X0,X1) | ? [X2,X3] : (X2 != X3 & r2_hidden(k4_tarski(X3,X2),X0) & r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(X3,X1) & r2_hidden(X2,X1))) & (! [X4,X5] : (X4 = X5 | ~r2_hidden(k4_tarski(X5,X4),X0) | ~r2_hidden(k4_tarski(X4,X5),X0) | ~r2_hidden(X5,X1) | ~r2_hidden(X4,X1)) | ~r4_relat_2(X0,X1))) | ~v1_relat_1(X0)) [rectify 141]
143. ! [X1,X0] : (? [X2,X3] : (X2 != X3 & r2_hidden(k4_tarski(X3,X2),X0) & r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(X3,X1) & r2_hidden(X2,X1)) => (sK3(X0,X1) != sK4(X0,X1) & r2_hidden(k4_tarski(sK4(X0,X1),sK3(X0,X1)),X0) & r2_hidden(k4_tarski(sK3(X0,X1),sK4(X0,X1)),X0) & r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK3(X0,X1),X1))) [choice axiom]
144. ! [X0] : (! [X1] : ((r4_relat_2(X0,X1) | (sK3(X0,X1) != sK4(X0,X1) & r2_hidden(k4_tarski(sK4(X0,X1),sK3(X0,X1)),X0) & r2_hidden(k4_tarski(sK3(X0,X1),sK4(X0,X1)),X0) & r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK3(X0,X1),X1))) & (! [X4,X5] : (X4 = X5 | ~r2_hidden(k4_tarski(X5,X4),X0) | ~r2_hidden(k4_tarski(X4,X5),X0) | ~r2_hidden(X5,X1) | ~r2_hidden(X4,X1)) | ~r4_relat_2(X0,X1))) | ~v1_relat_1(X0)) [skolemisation 142,143]
163. v1_relat_1(sK0) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
164. X3 = X4 | ~r2_hidden(k4_tarski(X4,X3),sK0) | ~r2_hidden(k4_tarski(X3,X4),sK0) | v4_relat_2(sK0) [cnf transformation 139] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
165. r2_hidden(k4_tarski(sK1,sK2),sK0) | ~v4_relat_2(sK0) [cnf transformation 139] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
166. r2_hidden(k4_tarski(sK2,sK1),sK0) | ~v4_relat_2(sK0) [cnf transformation 139] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
167. sK1 != sK2 | ~v4_relat_2(sK0) [cnf transformation 139] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
180. r4_relat_2(X0,k1_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 140] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
181. ~r4_relat_2(X0,k1_relat_1(X0)) | v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 140] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
182. ~r2_hidden(k4_tarski(X5,X4),X0) | X4 = X5 | ~r2_hidden(k4_tarski(X4,X5),X0) | ~r2_hidden(X5,X1) | ~r2_hidden(X4,X1) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 144] {a:0,w:24,wCS:9600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
185. r2_hidden(k4_tarski(sK3(X0,X1),sK4(X0,X1)),X0) | r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 144] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
186. r2_hidden(k4_tarski(sK4(X0,X1),sK3(X0,X1)),X0) | r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 144] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
187. sK3(X0,X1) != sK4(X0,X1) | r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 144] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
214. ~r2_hidden(k4_tarski(X0,X1),X2) | r2_hidden(X0,k1_relat_1(X2)) | ~v1_relat_1(X2) [cnf transformation 128] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
215. ~r2_hidden(k4_tarski(X0,X1),X2) | r2_hidden(X1,k1_relat_1(X2)) | ~v1_relat_1(X2) [cnf transformation 128] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
244. 1 <=> v4_relat_2(sK0) [avatar definition]
245. v4_relat_2(sK0) <- (1) [avatar component clause 244] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
248. 2 <=> sK1 = sK2 [avatar definition]
250. sK1 != sK2 <- (~2) [avatar component clause 248] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
251. ~1 | ~2 [avatar split clause 167,248,244]
253. 3 <=> r2_hidden(k4_tarski(sK2,sK1),sK0) [avatar definition]
255. r2_hidden(k4_tarski(sK2,sK1),sK0) <- (3) [avatar component clause 253] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
256. ~1 | 3 [avatar split clause 166,253,244]
258. 4 <=> r2_hidden(k4_tarski(sK1,sK2),sK0) [avatar definition]
260. r2_hidden(k4_tarski(sK1,sK2),sK0) <- (4) [avatar component clause 258] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
261. ~1 | 4 [avatar split clause 165,258,244]
263. 5 <=> ! [X3,X4] : (X3 = X4 | ~r2_hidden(k4_tarski(X3,X4),sK0) | ~r2_hidden(k4_tarski(X4,X3),sK0)) [avatar definition]
264. ~r2_hidden(k4_tarski(X4,X3),sK0) | ~r2_hidden(k4_tarski(X3,X4),sK0) | X3 = X4 <- (5) [avatar component clause 263] {a:0,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
265. 1 | 5 [avatar split clause 164,263,244]
267. 6 <=> v1_relat_1(sK0) [avatar definition]
269. v1_relat_1(sK0) <- (6) [avatar component clause 267] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
270. 6 [avatar split clause 163,267]
488. r4_relat_2(sK0,X0) | ~v1_relat_1(sK0) | ~r2_hidden(k4_tarski(sK4(sK0,X0),sK3(sK0,X0)),sK0) | sK4(sK0,X0) = sK3(sK0,X0) <- (5) [resolution 185,264] {a:1,w:21,wCS:2100,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
496. ~r2_hidden(k4_tarski(sK4(sK0,X0),sK3(sK0,X0)),sK0) | r4_relat_2(sK0,X0) | sK4(sK0,X0) = sK3(sK0,X0) <- (5, 6) [subsumption resolution 488,269] {a:1,w:19,wCS:1900,nSel:1,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
498. r4_relat_2(sK0,X0) | ~v1_relat_1(sK0) | r4_relat_2(sK0,X0) | sK4(sK0,X0) = sK3(sK0,X0) <- (5, 6) [resolution 186,496] {a:2,w:15,wCS:1500,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
507. r4_relat_2(sK0,X0) | ~v1_relat_1(sK0) | sK4(sK0,X0) = sK3(sK0,X0) <- (5, 6) [duplicate literal removal 498] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
508. sK4(sK0,X0) = sK3(sK0,X0) | r4_relat_2(sK0,X0) <- (5, 6) [subsumption resolution 507,269] {a:2,w:10,wCS:1000,nSel:1,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
521. sK3(sK0,X2) != sK3(sK0,X2) | r4_relat_2(sK0,X2) | ~v1_relat_1(sK0) | r4_relat_2(sK0,X2) <- (5, 6) [superposition 187,508] {a:3,w:15,wCS:1500,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
524. sK3(sK0,X2) != sK3(sK0,X2) | r4_relat_2(sK0,X2) | ~v1_relat_1(sK0) <- (5, 6) [duplicate literal removal 521] {a:3,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
525. r4_relat_2(sK0,X2) | ~v1_relat_1(sK0) <- (5, 6) [trivial inequality removal 524] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
530. r4_relat_2(sK0,X2) <- (5, 6) [subsumption resolution 525,269] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
532. v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (5, 6) [resolution 530,181] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
537. v4_relat_2(sK0) <- (5, 6) [subsumption resolution 532,269] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
538. 1 | ~5 | ~6 [avatar split clause 537,267,263,244]
557. 32 <=> r2_hidden(sK1,k1_relat_1(sK0)) [avatar definition]
559. r2_hidden(sK1,k1_relat_1(sK0)) <- (32) [avatar component clause 557] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
563. 33 <=> r2_hidden(sK2,k1_relat_1(sK0)) [avatar definition]
565. r2_hidden(sK2,k1_relat_1(sK0)) <- (33) [avatar component clause 563] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
584. 37 <=> ! [X0] : (~r2_hidden(sK2,X0) | ~r4_relat_2(sK0,X0) | ~r2_hidden(sK1,X0)) [avatar definition]
585. ~r4_relat_2(sK0,X0) | ~r2_hidden(sK2,X0) | ~r2_hidden(sK1,X0) <- (37) [avatar component clause 584] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
587. sK1 = sK2 | ~r2_hidden(k4_tarski(sK2,sK1),sK0) | ~r2_hidden(sK1,X0) | ~r2_hidden(sK2,X0) | ~r4_relat_2(sK0,X0) | ~v1_relat_1(sK0) <- (4) [resolution 260,182] {a:1,w:19,wCS:1900,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
588. r2_hidden(sK2,k1_relat_1(sK0)) | ~v1_relat_1(sK0) <- (4) [resolution 260,215] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
589. r2_hidden(sK1,k1_relat_1(sK0)) | ~v1_relat_1(sK0) <- (4) [resolution 260,214] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
593. ~r2_hidden(k4_tarski(sK2,sK1),sK0) | ~r2_hidden(sK1,X0) | ~r2_hidden(sK2,X0) | ~r4_relat_2(sK0,X0) | ~v1_relat_1(sK0) <- (~2, 4) [subsumption resolution 587,250] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
594. ~r2_hidden(sK1,X0) | ~r2_hidden(sK2,X0) | ~r4_relat_2(sK0,X0) | ~v1_relat_1(sK0) <- (~2, 3, 4) [subsumption resolution 593,255] {a:1,w:11,wCS:1100,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
595. ~r2_hidden(sK1,X0) | ~r2_hidden(sK2,X0) | ~r4_relat_2(sK0,X0) <- (~2, 3, 4, 6) [subsumption resolution 594,269] {a:1,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
596. 37 | 2 | ~3 | ~4 | ~6 [avatar split clause 595,267,258,253,248,584]
597. r2_hidden(sK2,k1_relat_1(sK0)) <- (4, 6) [subsumption resolution 588,269] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
598. 33 | ~4 | ~6 [avatar split clause 597,267,258,563]
599. r2_hidden(sK1,k1_relat_1(sK0)) <- (4, 6) [subsumption resolution 589,269] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
600. 32 | ~4 | ~6 [avatar split clause 599,267,258,557]
681. ~r2_hidden(sK2,k1_relat_1(sK0)) | ~r2_hidden(sK1,k1_relat_1(sK0)) | ~v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (37) [resolution 585,180] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
683. ~r2_hidden(sK1,k1_relat_1(sK0)) | ~v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (33, 37) [subsumption resolution 681,565] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
684. ~v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (32, 33, 37) [subsumption resolution 683,559] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
685. ~v1_relat_1(sK0) <- (1, 32, 33, 37) [subsumption resolution 684,245] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
686. $false <- (1, 6, 32, 33, 37) [subsumption resolution 685,269] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
687. ~1 | ~6 | ~32 | ~33 | ~37 [avatar contradiction clause 686]
688. $false [avatar sat refutation 251,256,261,265,270,538,596,598,600,687] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:325,thDist:-325}
% SZS output end Proof for wellord1__l3_wellord1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 62
% >>> Preprocessing
% Introduced skolems: 14
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 80
% Generated clauses: 319
% Active clauses: 107
% Passive clauses: 157
% Final active clauses: 101
% Final passive clauses: 40
% Main loop iterations started: 115
% 
% >>> Simplifying Inferences
% Duplicate literals: 18
% Trivial inequalities: 1
% Fw subsumption resolutions: 33
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 22
% Forward subsumptions: 40
% Backward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 80
% Unit resulting resolution: 1
% Forward superposition: 30
% Backward superposition: 26
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 10
% Split components: 20
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 59
% SAT solver unit clauses: 20
% SAT solver binary clauses: 20
% 
% Memory used [KB]: 5373
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% other: 0.004 s ( own 0.003 s ) 

