i: [1,56,0,2,1,1,0,0] t87_funct_4
a: [3,95,0,2,1,56]
s: 95
i: [1,55,0,2,1,1,0,0] t87_funct_4
a: [3,100,0,2,1,55]
s: 100
i: [1,54,0,2,1,1,0,0] t87_funct_4
a: [3,105,0,2,1,54]
s: 105
i: [1,53,0,2,1,1,0,0] t87_funct_4
a: [3,110,0,2,1,53]
s: 110
i: [1,52,0,2,1,1,0,0] t87_funct_4
a: [3,115,0,2,1,52]
s: 115
i: [1,51,0,2,1,1,0,0] t87_funct_4
a: [3,120,0,2,1,51]
s: 120
i: [1,58,0,3,1,1,0,0] t87_funct_4
a: [3,85,0,3,1,58]
s: 85
i: [1,57,0,3,1,1,0,0] t87_funct_4
a: [3,90,0,3,1,57]
s: 90
i: [1,59,0,5,1,1,0,0] t87_funct_4
a: [3,80,0,5,1,59]
s: 80
i: [1,76,0,2,1,0,0,1] rc1_funct_1
a: [3,132,0,2,1,76]
s: 132
i: [1,75,0,2,1,0,0,1] rc1_funct_1
a: [3,137,0,2,1,75]
s: 137
i: [1,68,0,3,1,0,0,1] reflexivity_r1_tarski
s: 68
i: [1,67,0,4,1,0,0,255] existence_m1_subset_1
s: 67
i: [1,69,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 69
i: [1,63,0,7,2,0,0,255] t3_subset
s: 63
i: [1,64,0,7,2,0,0,255] t3_subset
s: 64
d: [2,139,1,5,1,0,40,63,67]
s: 139
i: [1,70,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 70
i: [1,60,0,8,3,0,0,255] cc2_relat_1
s: 60
i: [1,74,0,8,3,0,0,2] fc3_relat_1
s: 74
i: [1,65,0,9,3,0,0,1] t1_xboole_1
s: 65
d: [2,147,1,6,2,1,40,65,85]
s: 147
i: [1,71,0,13,5,0,0,255] idempotence_k1_funct_4
a: [3,127,0,9,3,71]
s: 127
d: [2,152,2,5,1,1,40,147,139]
a: [3,156,2,5,1,152]
s: 156
i: [1,61,0,10,4,0,0,255] cc3_funct_1
s: 61
i: [1,66,0,11,3,0,0,2] t8_xboole_1
s: 66
d: [2,148,1,6,2,1,40,65,90]
s: 148
i: [1,72,0,12,5,0,0,1] dt_k1_funct_4
s: 72
d: [2,168,2,5,1,1,40,148,139]
a: [3,173,2,5,1,168]
s: 173
i: [1,73,0,12,5,0,0,1] dt_k1_funct_4
s: 73
i: [1,62,0,15,5,0,0,1] t29_funct_4
s: 62
d: [2,158,3,8,2,1,40,156,65]
s: 158
d: [2,141,1,6,2,0,40,60,67]
s: 141
d: [2,189,4,7,1,1,40,158,139]
a: [3,194,4,7,1,189]
s: 194
d: [2,142,1,7,3,0,40,60,64]
s: 142
d: [2,159,1,8,3,0,40,61,67]
s: 159
d: [2,203,3,6,2,1,40,142,156]
d: [2,206,3,4,1,2,21,203,100]
a: [3,210,3,4,1,206]
s: 210
d: [2,160,1,9,4,0,40,61,64]
s: 160
d: [2,204,3,6,2,1,40,142,173]
d: [2,212,3,4,1,2,21,204,100]
a: [3,216,3,4,1,212]
s: 216
d: [2,182,1,13,5,0,44,62,69]
d: [2,185,1,9,3,0,18,182]
s: 185
d: [2,161,1,14,4,0,40,66,65]
s: 161
d: [2,230,3,8,3,1,40,160,156]
d: [2,233,3,6,2,2,21,230,95]
d: [2,234,3,4,1,3,21,233,100]
a: [3,238,3,4,1,234]
s: 238
d: [2,183,1,15,5,0,44,62,70]
s: 183
d: [2,231,3,8,3,1,40,160,173]
d: [2,240,3,6,2,2,21,231,95]
d: [2,241,3,4,1,3,21,240,100]
a: [3,245,3,4,1,241]
s: 245
d: [2,184,1,15,5,0,44,62,70]
s: 184
d: [2,180,1,18,6,0,40,62,65]
s: 180
d: [2,257,2,9,3,1,40,185,148]
d: [2,264,2,7,2,2,21,257,115]
d: [2,265,2,5,1,3,21,264,120]
a: [3,269,2,5,1,265]
s: 269
d: [2,258,2,9,3,1,40,185,147]
d: [2,271,2,7,2,2,21,258,105]
d: [2,272,2,5,1,3,21,271,110]
a: [3,276,2,5,1,272]
s: 276
d: [2,313,3,8,3,1,40,269,160]
d: [2,317,3,6,2,2,21,313,95]
d: [2,318,3,4,1,3,21,317,100]
a: [3,322,3,4,1,318]
s: 322
d: [2,162,2,11,3,1,40,66,147]
s: 162
d: [2,169,2,11,3,1,40,148,66]
s: 169
d: [2,314,3,6,2,1,40,269,142]
d: [2,324,3,4,1,2,21,314,100]
a: [3,328,3,4,1,324]
s: 328
d: [2,150,2,8,2,0,40,65,139]
s: 150
d: [2,330,3,8,3,1,40,276,160]
d: [2,334,3,6,2,2,21,330,95]
d: [2,335,3,4,1,3,21,334,100]
a: [3,339,3,4,1,335]
s: 339
d: [2,255,2,10,4,0,40,185,142]
d: [2,262,2,8,3,0,18,255]
s: 262
d: [2,254,2,12,5,0,40,185,160]
d: [2,263,2,8,3,0,18,254]
s: 263
d: [2,331,3,6,2,1,40,276,142]
d: [2,341,3,4,1,2,21,331,100]
a: [3,345,3,4,1,341]
s: 345
d: [2,199,2,12,4,0,40,142,66]
s: 199
d: [2,205,5,8,2,1,40,142,194]
d: [2,218,5,6,1,2,21,205,100]
a: [3,222,5,6,1,218]
s: 222
d: [2,256,2,12,4,0,40,185,65]
s: 256
d: [2,286,2,13,3,0,40,161,139]
s: 286
d: [2,232,5,10,3,1,40,160,194]
d: [2,247,5,8,2,2,21,232,95]
d: [2,248,5,6,1,3,21,247,100]
a: [3,252,5,6,1,248]
s: 252
d: [2,226,2,14,5,0,40,160,66]
s: 226
d: [2,177,3,8,2,1,40,173,65]
s: 177
d: [2,305,2,17,5,0,40,180,139]
s: 305
d: [2,296,2,18,6,0,40,183,65]
s: 296
d: [2,410,4,7,1,1,40,177,139]
a: [3,417,4,7,1,410]
s: 417
d: [2,284,2,19,7,0,40,161,62]
s: 284
d: [2,465,3,0,0,7,52,120,115,110,105,85,90,80,284]
e: 465
d: [2,466,3,14,6,1,40,284,80]
d: [2,482,3,11,5,2,21,466,90]
d: [2,483,3,8,4,3,21,482,85]
d: [2,484,3,6,3,4,21,483,105]
d: [2,485,3,4,2,5,21,484,110]
d: [2,486,3,2,1,6,21,485,115]
d: [2,487,3,0,0,7,21,486,120]
e: 487
f: 487,51,52,53,54,57,58,59
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t87_funct_4
% SZS output start Proof for funct_4__t87_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_tarski(X1,X2) & r1_tarski(X0,X2)) => r1_tarski(k1_funct_4(X0,X1),X2))))) [input t87_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_tarski(X1,X2) & r1_tarski(X0,X2)) => r1_tarski(k1_funct_4(X0,X1),X2))))) [negated conjecture 1]
16. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
17. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => r1_tarski(k1_funct_4(X0,X1),k2_xboole_0(X0,X1)))) [input t29_funct_4]
19. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [input t8_xboole_1]
21. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 16]
22. ! [X0,X1,X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [flattening 19]
25. ? [X0] : (? [X1] : (? [X2] : ((~r1_tarski(k1_funct_4(X0,X1),X2) & (r1_tarski(X1,X2) & r1_tarski(X0,X2))) & (v1_funct_1(X2) & v1_relat_1(X2))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
26. ? [X0] : (? [X1] : (? [X2] : (~r1_tarski(k1_funct_4(X0,X1),X2) & r1_tarski(X1,X2) & r1_tarski(X0,X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 25]
30. ! [X0] : (! [X1] : (r1_tarski(k1_funct_4(X0,X1),k2_xboole_0(X0,X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 17]
31. ! [X0] : (! [X1] : (r1_tarski(k1_funct_4(X0,X1),k2_xboole_0(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 30]
32. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 21]
33. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 32]
34. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | (~r1_tarski(X2,X1) | ~r1_tarski(X0,X1))) [ennf transformation 22]
35. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1)) [flattening 34]
42. ? [X0] : (? [X1] : (? [X2] : (~r1_tarski(k1_funct_4(X0,X1),X2) & r1_tarski(X1,X2) & r1_tarski(X0,X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (~r1_tarski(k1_funct_4(sK0,X1),X2) & r1_tarski(X1,X2) & r1_tarski(sK0,X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
43. ? [X1] : (? [X2] : (~r1_tarski(k1_funct_4(sK0,X1),X2) & r1_tarski(X1,X2) & r1_tarski(sK0,X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) => (? [X2] : (~r1_tarski(k1_funct_4(sK0,sK1),X2) & r1_tarski(sK1,X2) & r1_tarski(sK0,X2) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
44. ? [X2] : (~r1_tarski(k1_funct_4(sK0,sK1),X2) & r1_tarski(sK1,X2) & r1_tarski(sK0,X2) & v1_funct_1(X2) & v1_relat_1(X2)) => (~r1_tarski(k1_funct_4(sK0,sK1),sK2) & r1_tarski(sK1,sK2) & r1_tarski(sK0,sK2) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
45. ((~r1_tarski(k1_funct_4(sK0,sK1),sK2) & r1_tarski(sK1,sK2) & r1_tarski(sK0,sK2) & v1_funct_1(sK2) & v1_relat_1(sK2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 26,44,43,42]
51. v1_relat_1(sK0) [cnf transformation 45] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
52. v1_funct_1(sK0) [cnf transformation 45] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
53. v1_relat_1(sK1) [cnf transformation 45] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
54. v1_funct_1(sK1) [cnf transformation 45] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
57. r1_tarski(sK0,sK2) [cnf transformation 45] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
58. r1_tarski(sK1,sK2) [cnf transformation 45] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
59. ~r1_tarski(k1_funct_4(sK0,sK1),sK2) [cnf transformation 45] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
62. r1_tarski(k1_funct_4(X0,X1),k2_xboole_0(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 33] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
66. r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1) [cnf transformation 35] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
78. 1 <=> r1_tarski(k1_funct_4(sK0,sK1),sK2) [avatar definition]
80. ~r1_tarski(k1_funct_4(sK0,sK1),sK2) <- (~1) [avatar component clause 78] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
81. ~1 [avatar split clause 59,78]
83. 2 <=> r1_tarski(sK1,sK2) [avatar definition]
85. r1_tarski(sK1,sK2) <- (2) [avatar component clause 83] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
86. 2 [avatar split clause 58,83]
88. 3 <=> r1_tarski(sK0,sK2) [avatar definition]
90. r1_tarski(sK0,sK2) <- (3) [avatar component clause 88] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
91. 3 [avatar split clause 57,88]
103. 6 <=> v1_funct_1(sK1) [avatar definition]
105. v1_funct_1(sK1) <- (6) [avatar component clause 103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
106. 6 [avatar split clause 54,103]
108. 7 <=> v1_relat_1(sK1) [avatar definition]
110. v1_relat_1(sK1) <- (7) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
111. 7 [avatar split clause 53,108]
113. 8 <=> v1_funct_1(sK0) [avatar definition]
115. v1_funct_1(sK0) <- (8) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
116. 8 [avatar split clause 52,113]
118. 9 <=> v1_relat_1(sK0) [avatar definition]
120. v1_relat_1(sK0) <- (9) [avatar component clause 118] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
121. 9 [avatar split clause 51,118]
161. ~r1_tarski(X3,k2_xboole_0(X2,X0)) | ~r1_tarski(X2,X1) | r1_tarski(X3,X1) | ~r1_tarski(X0,X1) [resolution 66,65] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
284. r1_tarski(k1_funct_4(X0,X2),X1) | ~r1_tarski(X0,X1) | ~r1_tarski(X2,X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [resolution 161,62] {a:2,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
466. ~r1_tarski(sK0,sK2) | ~r1_tarski(sK1,sK2) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [resolution 284,80] {a:3,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
482. ~r1_tarski(sK1,sK2) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 3) [subsumption resolution 466,90] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
483. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3) [subsumption resolution 482,85] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
484. ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 6) [subsumption resolution 483,105] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
485. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 6, 7) [subsumption resolution 484,110] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
486. ~v1_relat_1(sK0) <- (~1, 2, 3, 6, 7, 8) [subsumption resolution 485,115] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
487. $false <- (~1, 2, 3, 6, 7, 8, 9) [subsumption resolution 486,120] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
488. 1 | ~2 | ~3 | ~6 | ~7 | ~8 | ~9 [avatar contradiction clause 487]
489. $false [avatar sat refutation 81,86,91,106,111,116,121,488] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:366,thDist:-366}
% SZS output end Proof for funct_4__t87_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 306
% Active clauses: 67
% Passive clauses: 118
% Final active clauses: 67
% Final passive clauses: 50
% Main loop iterations started: 68
% 
% >>> Simplifying Inferences
% Duplicate literals: 58
% Fw subsumption resolutions: 45
% Fw demodulations: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Forward subsumptions: 52
% 
% >>> Generating Inferences
% Binary resolution: 121
% Unit resulting resolution: 1
% Forward superposition: 51
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 36
% SAT solver unit clauses: 11
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9722
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

