i: [1,301,0,27,8,1,0,0] t2_gate_3
a: [3,316,0,2,1,301]
s: 316
a: [3,324,0,2,1,301]
s: 324
a: [3,332,0,2,1,301]
s: 332
a: [3,340,0,2,1,301]
s: 340
i: [1,45,0,6,2,1,0,0] t2_gate_3
a: [3,365,0,2,1,45]
s: 365
i: [1,43,0,7,2,1,0,0] t2_gate_3
a: [3,375,0,2,1,43]
s: 375
i: [1,33,0,6,2,1,0,0] t2_gate_3
a: [3,409,0,2,1,33]
s: 409
i: [1,31,0,7,2,1,0,0] t2_gate_3
a: [3,419,0,2,1,31]
s: 419
i: [1,29,0,7,2,1,0,0] t2_gate_3
a: [3,429,0,2,1,29]
s: 429
i: [1,27,0,8,2,1,0,0] t2_gate_3
a: [3,439,0,2,1,27]
s: 439
a: [3,321,0,4,1,301]
s: 321
a: [3,360,0,4,1,45]
s: 360
i: [1,41,0,6,2,1,0,0] t2_gate_3
a: [3,380,0,4,1,41]
s: 380
a: [3,404,0,4,1,33]
s: 404
a: [3,370,0,5,1,43]
s: 370
i: [1,39,0,7,2,1,0,0] t2_gate_3
a: [3,386,0,5,1,39]
s: 386
i: [1,37,0,7,2,1,0,0] t2_gate_3
a: [3,392,0,5,1,37]
s: 392
a: [3,414,0,5,1,31]
s: 414
a: [3,424,0,5,1,29]
s: 424
i: [1,35,0,8,2,1,0,0] t2_gate_3
a: [3,398,0,6,1,35]
s: 398
a: [3,434,0,6,1,27]
s: 434
i: [1,302,0,5,2,0,0,1] fc2_gate_1
s: 302
i: [1,303,0,5,2,0,0,1] fc1_gate_1
s: 303
i: [1,306,0,6,2,0,0,1] fc5_gate_1
s: 306
i: [1,307,0,6,2,0,0,1] fc4_gate_1
s: 307
i: [1,304,0,7,1,0,0,255] commutativity_k3_gate_1
s: 304
d: [2,448,1,2,1,1,40,307,360]
a: [3,469,1,2,1,448]
s: 469
i: [1,305,0,7,1,0,0,255] commutativity_k2_gate_1
s: 305
r: 398
r: 434
r: 404
r: 424
r: 380
r: 392
r: 370
d: [2,449,1,2,1,1,40,307,404]
a: [3,474,1,2,1,449]
s: 474
d: [2,481,0,4,1,1,23,404,305]
a: [3,501,0,4,1,481]
s: 501
d: [2,483,0,4,1,1,23,380,305]
a: [3,511,0,4,1,483]
s: 511
d: [2,443,1,3,1,1,40,307,398]
a: [3,456,1,3,1,443]
s: 456
d: [2,525,2,0,0,2,52,365,456,302]
e: 525
d: [2,526,2,2,1,1,40,456,302]
d: [2,528,2,0,0,2,21,526,365]
e: 528
r: 360
r: 365
r: 324
r: 332
r: 386
a: [3,381,0,4,1,41]
d: [2,533,0,4,1,1,22,381,305]
d: [2,534,0,0,0,2,21,533,511]
e: 534
r: 511
a: [3,325,0,2,1,301]
s: 325
a: [3,333,0,2,1,301]
s: 333
a: [3,364,0,2,1,45]
s: 364
a: [3,328,0,4,1,301]
s: 328
d: [2,444,1,3,1,1,40,307,434]
a: [3,461,1,3,1,444]
s: 461
d: [2,538,2,0,0,2,52,474,461,302]
e: 538
d: [2,539,2,2,1,1,40,461,302]
d: [2,541,2,0,0,2,21,539,474]
e: 541
r: 375
r: 429
r: 439
r: 461
a: [3,374,0,2,1,43]
s: 374
a: [3,428,0,2,1,29]
s: 428
a: [3,438,0,2,1,27]
s: 438
a: [3,361,0,4,1,45]
s: 361
a: [3,510,0,4,1,483]
s: 510
a: [3,387,0,5,1,39]
s: 387
d: [2,482,0,5,1,1,23,424,305]
a: [3,505,0,5,1,482]
s: 505
d: [2,484,0,5,1,1,23,392,305]
a: [3,516,0,5,1,484]
s: 516
d: [2,549,1,0,0,2,52,374,516,307]
e: 549
d: [2,550,1,2,1,1,40,516,307]
d: [2,552,1,0,0,2,21,550,374]
e: 552
r: 316
r: 516
a: [3,317,0,2,1,301]
s: 317
a: [3,515,0,5,1,484]
s: 515
d: [2,485,0,5,1,1,23,370,305]
a: [3,520,0,5,1,485]
s: 520
d: [2,479,0,6,1,1,23,398,305]
a: [3,491,0,6,1,479]
s: 491
d: [2,480,0,6,1,1,23,434,305]
a: [3,495,0,6,1,480]
s: 495
i: [1,308,0,8,3,0,0,1] fc3_gate_1
s: 308
r: 340
r: 361
r: 364
r: 456
a: [3,399,0,6,1,35]
d: [2,588,0,6,1,1,22,399,305]
d: [2,589,0,0,0,2,21,588,491]
e: 589
r: 333
r: 387
r: 491
s: 332
a: [3,341,0,2,1,301]
s: 341
s: 365
s: 360
s: 386
a: [3,455,1,3,1,443]
s: 455
a: [3,490,0,6,1,479]
s: 490
a: [3,460,1,3,1,444]
s: 460
a: [3,336,0,7,1,301]
d: [2,592,0,7,1,1,22,336,304]
a: [3,596,0,7,1,592]
s: 596
i: [1,309,0,8,3,0,0,1] fc6_gate_1
s: 309
d: [2,556,1,3,1,1,40,491,307]
a: [3,560,1,3,1,556]
s: 560
d: [2,604,1,7,2,1,40,309,596]
a: [3,610,1,3,1,604]
s: 610
d: [2,565,1,4,2,1,40,308,361]
d: [2,578,1,2,1,2,21,565,469]
a: [3,582,1,2,1,578]
s: 582
d: [2,477,1,6,2,0,44,306,304]
s: 477
d: [2,478,1,6,2,0,44,306,304]
s: 478
d: [2,486,1,6,2,0,44,307,305]
s: 486
d: [2,622,2,0,0,2,52,582,328,486]
e: 622
d: [2,623,2,0,0,2,52,582,360,486]
e: 623
d: [2,625,2,2,1,1,40,486,328]
d: [2,634,2,0,0,2,21,625,582]
e: 634
d: [2,627,2,2,1,1,40,486,360]
d: [2,636,2,0,0,2,21,627,582]
e: 636
r: 328
r: 332
r: 360
r: 365
r: 386
r: 455
r: 560
s: 333
s: 364
s: 361
s: 387
s: 456
a: [3,559,1,3,1,556]
s: 559
d: [2,646,2,0,0,2,52,374,559,303]
e: 646
d: [2,647,2,2,1,1,40,559,303]
d: [2,649,2,0,0,2,21,647,374]
e: 649
r: 333
r: 361
r: 364
r: 387
r: 456
r: 559
r: 582
r: 610
s: 332
s: 365
s: 360
a: [3,581,1,2,1,578]
s: 581
s: 386
s: 455
s: 560
d: [2,563,1,6,2,1,40,308,495]
a: [3,575,1,3,1,563]
s: 575
d: [2,659,2,0,0,2,52,469,575,302]
e: 659
d: [2,660,2,2,1,1,40,575,302]
d: [2,662,2,0,0,2,21,660,469]
e: 662
r: 414
r: 419
r: 575
a: [3,418,0,2,1,31]
s: 418
a: [3,415,0,5,1,31]
s: 415
a: [3,576,1,3,1,563]
s: 576
a: [3,609,1,3,1,604]
s: 609
a: [3,614,1,4,1,604]
s: 614
d: [2,669,2,0,0,2,52,418,614,307]
e: 669
d: [2,671,2,2,1,1,40,614,307]
d: [2,673,2,0,0,2,21,671,418]
e: 673
r: 332
r: 341
r: 360
r: 365
r: 386
r: 455
r: 490
r: 581
r: 609
r: 614
s: 333
s: 340
s: 364
a: [3,312,0,4,1,301]
s: 312
d: [2,677,1,0,0,2,52,438,312,307]
e: 677
d: [2,679,1,2,1,1,40,312,307]
d: [2,681,1,0,0,2,21,679,438]
e: 681
r: 409
r: 438
r: 460
r: 469
r: 474
r: 495
a: [3,405,0,4,1,33]
d: [2,690,0,4,1,1,22,405,305]
d: [2,691,0,0,0,2,21,690,501]
e: 691
r: 428
r: 501
r: 505
r: 312
a: [3,337,0,7,1,301]
d: [2,698,0,7,1,1,22,337,304]
d: [2,699,0,0,0,2,21,698,596]
e: 699
r: 596
a: [3,408,0,2,1,33]
s: 408
s: 429
s: 439
s: 469
a: [3,313,0,4,1,301]
s: 313
s: 361
s: 582
a: [3,500,0,4,1,481]
s: 500
a: [3,473,1,2,1,449]
s: 473
s: 387
a: [3,506,0,5,1,482]
s: 506
s: 456
s: 491
s: 461
a: [3,496,0,6,1,480]
s: 496
d: [2,719,2,0,0,2,52,576,496,307]
e: 719
d: [2,721,1,3,1,1,40,496,307]
d: [2,723,1,0,0,2,21,721,576]
e: 723
r: 439
r: 496
s: 438
s: 610
s: 495
a: [3,595,0,7,1,592]
s: 595
d: [2,729,2,0,0,2,52,610,595,477]
e: 729
d: [2,730,2,3,1,1,40,595,477]
d: [2,733,2,0,0,2,21,730,610]
e: 733
r: 340
r: 374
r: 491
r: 520
r: 560
r: 595
r: 321
r: 429
r: 461
r: 506
s: 340
s: 374
s: 428
a: [3,320,0,4,1,301]
s: 320
d: [2,749,2,0,0,2,52,582,320,486]
e: 749
d: [2,750,1,0,0,2,52,428,320,307]
e: 750
d: [2,751,2,2,1,1,40,320,486]
d: [2,755,2,0,0,2,21,751,582]
e: 755
d: [2,752,1,2,1,1,40,320,307]
d: [2,757,1,0,0,2,21,752,428]
e: 757
r: 317
r: 320
r: 374
r: 428
r: 515
s: 316
s: 375
s: 429
s: 461
s: 506
s: 560
d: [2,768,2,0,0,2,52,375,560,302]
e: 768
d: [2,769,2,2,1,1,40,560,302]
d: [2,771,2,0,0,2,21,769,375]
e: 771
r: 340
r: 560
s: 341
s: 516
s: 559
a: [3,521,0,5,1,485]
s: 521
d: [2,778,2,0,0,2,52,582,521,307]
e: 778
d: [2,780,1,2,1,1,40,521,307]
d: [2,782,1,0,0,2,21,780,582]
e: 782
r: 316
r: 341
r: 375
r: 415
r: 418
r: 429
r: 461
r: 469
r: 473
r: 506
r: 516
r: 521
r: 559
r: 576
r: 582
r: 610
s: 317
s: 340
s: 374
s: 419
s: 428
a: [3,468,1,2,1,448]
s: 468
s: 328
d: [2,795,1,0,0,2,52,408,328,307]
e: 795
d: [2,797,1,2,1,1,40,328,307]
d: [2,799,1,0,0,2,21,797,408]
e: 799
r: 328
s: 414
s: 474
s: 505
s: 460
s: 515
s: 520
s: 560
s: 491
s: 575
s: 595
s: 609
s: 581
a: [3,613,1,4,1,604]
s: 613
d: [2,821,2,0,0,3,52,419,581,613,308]
e: 821
d: [2,822,2,4,2,1,40,613,308]
d: [2,824,2,2,1,2,21,822,581]
d: [2,825,2,0,0,3,21,824,419]
e: 825
r: 414
r: 419
r: 474
r: 333
r: 340
r: 361
r: 364
r: 387
r: 408
r: 456
r: 468
r: 491
r: 500
r: 575
s: 332
s: 341
s: 365
s: 409
s: 418
s: 469
a: [3,329,0,4,1,301]
s: 329
d: [2,833,2,0,0,3,52,409,581,329,308]
e: 833
d: [2,834,1,4,2,1,40,329,308]
d: [2,836,1,2,1,2,21,834,581]
d: [2,837,1,0,0,3,21,836,409]
e: 837
r: 325
r: 329
r: 510
s: 324
s: 474
s: 321
s: 455
s: 360
s: 501
s: 576
s: 511
d: [2,848,1,0,0,2,52,374,511,307]
e: 848
d: [2,850,1,2,1,1,40,511,307]
d: [2,852,1,0,0,2,21,850,374]
e: 852
r: 321
r: 324
r: 332
r: 374
r: 409
r: 428
r: 460
r: 474
r: 501
r: 505
r: 511
r: 520
r: 560
s: 325
s: 333
s: 375
s: 408
s: 429
s: 473
s: 320
s: 500
s: 461
s: 510
d: [2,866,1,0,0,3,52,375,365,510,308]
e: 866
d: [2,867,1,4,2,1,40,510,308]
d: [2,869,1,2,1,2,21,867,365]
d: [2,870,1,0,0,3,21,869,375]
e: 870
r: 313
r: 317
r: 360
r: 365
r: 438
r: 455
r: 469
r: 495
r: 515
r: 576
s: 316
s: 364
s: 439
s: 468
s: 312
s: 456
s: 361
s: 387
s: 559
s: 415
s: 575
s: 506
d: [2,886,2,0,0,2,52,468,506,307]
e: 886
d: [2,888,1,2,1,1,40,506,307]
d: [2,890,1,0,0,2,21,888,468]
e: 890
i: [1,238,0,27,8,1,0,0] t2_gate_3
i: [1,94,0,27,8,1,0,0] t2_gate_3
i: [1,58,0,27,8,1,0,0] t2_gate_3
i: [1,49,0,27,8,1,0,0] t2_gate_3
i: [1,44,0,6,2,1,0,0] t2_gate_3
i: [1,42,0,7,2,1,0,0] t2_gate_3
i: [1,40,0,6,2,1,0,0] t2_gate_3
i: [1,38,0,7,2,1,0,0] t2_gate_3
i: [1,36,0,7,2,1,0,0] t2_gate_3
i: [1,34,0,8,2,1,0,0] t2_gate_3
i: [1,32,0,6,2,1,0,0] t2_gate_3
i: [1,30,0,7,2,1,0,0] t2_gate_3
i: [1,28,0,7,2,1,0,0] t2_gate_3
i: [1,26,0,8,2,1,0,0] t2_gate_3
d: [2,445,1,3,1,1,40,307,424]
d: [2,446,1,3,1,1,40,307,392]
d: [2,447,1,3,1,1,40,307,370]
d: [2,452,1,2,1,1,40,307,386]
d: [2,451,1,2,1,1,40,307,380]
d: [2,570,1,5,2,1,40,308,520]
d: [2,569,1,5,2,1,40,308,387]
d: [2,617,2,2,1,1,40,610,302]
d: [2,600,2,2,1,1,40,455,303]
d: [2,601,1,6,2,1,40,490,308]
d: [2,651,2,2,1,1,40,360,486]
d: [2,654,2,3,1,1,40,386,486]
d: [2,664,1,5,2,1,40,415,308]
d: [2,665,1,3,1,2,21,664,474]
d: [2,667,2,2,1,1,40,576,303]
d: [2,566,1,5,2,1,40,308,505]
d: [2,523,1,2,1,1,40,501,307]
d: [2,564,1,4,2,1,40,308,321]
d: [2,678,2,2,1,1,40,312,486]
d: [2,702,1,4,2,1,40,313,308]
d: [2,703,1,2,1,2,21,702,439]
a: [3,435,0,6,1,27]
d: [2,725,0,6,1,1,22,435,305]
d: [2,568,1,5,2,1,40,308,515]
d: [2,735,1,2,1,2,21,568,456]
d: [2,741,0,7,1,1,22,336,304]
d: [2,708,1,4,2,1,40,500,308]
d: [2,705,1,4,2,1,40,361,308]
d: [2,728,1,6,2,1,40,495,308]
a: [3,425,0,5,1,29]
d: [2,787,0,5,1,1,22,425,305]
d: [2,789,0,6,1,1,22,398,305]
a: [3,393,0,5,1,37]
d: [2,791,0,5,1,1,22,393,305]
a: [3,371,0,5,1,43]
d: [2,793,0,5,1,1,22,371,305]
d: [2,801,0,7,1,1,22,337,304]
d: [2,803,2,3,1,1,40,414,486]
d: [2,804,1,2,1,1,40,414,307]
d: [2,815,2,3,1,1,40,595,477]
d: [2,816,1,4,1,1,40,595,306]
d: [2,819,2,2,1,1,40,609,303]
d: [2,808,2,2,1,1,40,460,303]
d: [2,829,0,4,1,1,22,404,305]
d: [2,831,0,6,1,1,22,399,305]
d: [2,839,0,4,1,1,22,380,305]
d: [2,854,0,5,1,1,22,424,305]
d: [2,856,0,4,1,1,22,405,305]
d: [2,858,0,4,1,1,22,381,305]
d: [2,860,0,5,1,1,22,370,305]
d: [2,872,1,2,1,2,21,867,375]
d: [2,874,0,6,1,1,22,434,305]
d: [2,876,0,5,1,1,22,392,305]
f: 890,876,874,872,860,858,856,854,852,839,837,831,829,808,824,819,816,815,804,803,801,799,793,791,789,787,728,705,708,782,771,757,755,741,735,725,723,703,678,564,523,566,667,679,673,665,662,654,651,649,601,600,634,617,604,569,570,556,552,541,451,452,526,449,448,447,446,445,444,443,26,27,28,29,30,31,32,33,34,35,36,37,38,39,40,41,42,43,44,45,49,58,94,238
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_3__t2_gate_3
% SZS output start Proof for gate_3__t2_gate_3
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ~(~(~(v1_xboole_0(X4) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12)))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X3,X12))) & ~(v1_xboole_0(k2_gate_1(X3,X12)) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X1,X12))) & ~(v1_xboole_0(k2_gate_1(X1,X12)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(X0,X12))) & ~(v1_xboole_0(k2_gate_1(X0,X12)) & ~v1_xboole_0(X5))) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X8,X12))) & ~(v1_xboole_0(k2_gate_1(X8,X12)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X11,X10))) & ~(v1_xboole_0(k2_gate_1(X11,X10)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k2_gate_1(X9,X8))) & ~(v1_xboole_0(k2_gate_1(X9,X8)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) & ~v1_xboole_0(X0))) [input t2_gate_3]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ~(~(~(v1_xboole_0(X4) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12)))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X3,X12))) & ~(v1_xboole_0(k2_gate_1(X3,X12)) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X1,X12))) & ~(v1_xboole_0(k2_gate_1(X1,X12)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(X0,X12))) & ~(v1_xboole_0(k2_gate_1(X0,X12)) & ~v1_xboole_0(X5))) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X8,X12))) & ~(v1_xboole_0(k2_gate_1(X8,X12)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X11,X10))) & ~(v1_xboole_0(k2_gate_1(X11,X10)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k2_gate_1(X9,X8))) & ~(v1_xboole_0(k2_gate_1(X9,X8)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) & ~v1_xboole_0(X0))) [negated conjecture 1]
3. ! [X0,X1] : k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [input commutativity_k2_gate_1]
4. ! [X0,X1] : k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [input commutativity_k3_gate_1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. ! [X0] : (~v1_xboole_0(X0) => v1_xboole_0(k1_gate_1(X0))) [input fc2_gate_1]
10. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_gate_1(X0,X1))) [input fc3_gate_1]
11. ! [X0,X1] : (v1_xboole_0(X0) => v1_xboole_0(k2_gate_1(X0,X1))) [input fc4_gate_1]
12. ! [X0,X1] : (~v1_xboole_0(X1) => ~v1_xboole_0(k3_gate_1(X0,X1))) [input fc5_gate_1]
13. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k3_gate_1(X0,X1))) [input fc6_gate_1]
14. ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ~(~(~(v1_xboole_0(X4) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12)))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X3,X12))) & ~(v1_xboole_0(k2_gate_1(X3,X12)) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X1,X12))) & ~(v1_xboole_0(k2_gate_1(X1,X12)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(X0,X12))) & ~(v1_xboole_0(k2_gate_1(X0,X12)) & ~v1_xboole_0(X5))) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k2_gate_1(X8,X12))) & ~(v1_xboole_0(k2_gate_1(X8,X12)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12)) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X11,X10))) & ~(v1_xboole_0(k2_gate_1(X11,X10)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k2_gate_1(X9,X8))) & ~(v1_xboole_0(k2_gate_1(X9,X8)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & ~(v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) & ~v1_xboole_0(X0))) [flattening 2]
15. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (((v1_xboole_0(X4) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12)))) | (v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12))) & ~v1_xboole_0(X4)) | (v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X3,X12))) | (v1_xboole_0(k2_gate_1(X3,X12)) & ~v1_xboole_0(X6)) | (v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X1,X12))) | (v1_xboole_0(k2_gate_1(X1,X12)) & ~v1_xboole_0(X7)) | (v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(X0,X12))) | (v1_xboole_0(k2_gate_1(X0,X12)) & ~v1_xboole_0(X5))) & (~v1_xboole_0(X11) | v1_xboole_0(k2_gate_1(X8,X12))) & (~v1_xboole_0(k2_gate_1(X8,X12)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12)) | v1_xboole_0(X10)) & (~v1_xboole_0(X7) | v1_xboole_0(k2_gate_1(X11,X10))) & (~v1_xboole_0(k2_gate_1(X11,X10)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & (~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k2_gate_1(X9,X8))) & (~v1_xboole_0(k2_gate_1(X9,X8)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & (~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) | v1_xboole_0(X0))) [ennf transformation 14]
16. ! [X0] : (v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0)) [ennf transformation 9]
17. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
18. ! [X0,X1] : (~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1)) [ennf transformation 12]
19. ! [X0,X1] : (v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 11]
20. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 10]
21. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 20]
22. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 13]
23. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 22]
24. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (((v1_xboole_0(X4) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12)))) | (v1_xboole_0(k3_gate_1(k2_gate_1(X2,X12),k1_gate_1(X12))) & ~v1_xboole_0(X4)) | (v1_xboole_0(X6) & ~v1_xboole_0(k2_gate_1(X3,X12))) | (v1_xboole_0(k2_gate_1(X3,X12)) & ~v1_xboole_0(X6)) | (v1_xboole_0(X7) & ~v1_xboole_0(k2_gate_1(X1,X12))) | (v1_xboole_0(k2_gate_1(X1,X12)) & ~v1_xboole_0(X7)) | (v1_xboole_0(X5) & ~v1_xboole_0(k2_gate_1(X0,X12))) | (v1_xboole_0(k2_gate_1(X0,X12)) & ~v1_xboole_0(X5))) & (~v1_xboole_0(X11) | v1_xboole_0(k2_gate_1(X8,X12))) & (~v1_xboole_0(k2_gate_1(X8,X12)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X12)) | v1_xboole_0(X10)) & (~v1_xboole_0(X7) | v1_xboole_0(k2_gate_1(X11,X10))) & (~v1_xboole_0(k2_gate_1(X11,X10)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10)))) & (~v1_xboole_0(k2_gate_1(X11,k1_gate_1(X10))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X11),X10)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X11),k1_gate_1(X10))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k2_gate_1(X9,X8))) & (~v1_xboole_0(k2_gate_1(X9,X8)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8)))) & (~v1_xboole_0(k2_gate_1(X9,k1_gate_1(X8))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),X8)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(X9),k1_gate_1(X8))) | v1_xboole_0(X0))) => (((v1_xboole_0(sK4) & ~v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12)))) | (v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12))) & ~v1_xboole_0(sK4)) | (v1_xboole_0(sK6) & ~v1_xboole_0(k2_gate_1(sK3,sK12))) | (v1_xboole_0(k2_gate_1(sK3,sK12)) & ~v1_xboole_0(sK6)) | (v1_xboole_0(sK7) & ~v1_xboole_0(k2_gate_1(sK1,sK12))) | (v1_xboole_0(k2_gate_1(sK1,sK12)) & ~v1_xboole_0(sK7)) | (v1_xboole_0(sK5) & ~v1_xboole_0(k2_gate_1(sK0,sK12))) | (v1_xboole_0(k2_gate_1(sK0,sK12)) & ~v1_xboole_0(sK5))) & (~v1_xboole_0(sK11) | v1_xboole_0(k2_gate_1(sK8,sK12))) & (~v1_xboole_0(k2_gate_1(sK8,sK12)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK12))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK12)) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK7) | v1_xboole_0(k2_gate_1(sK11,sK10))) & (~v1_xboole_0(k2_gate_1(sK11,sK10)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10)))) & (~v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k2_gate_1(sK9,sK8))) & (~v1_xboole_0(k2_gate_1(sK9,sK8)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8)))) & (~v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) | v1_xboole_0(sK0))) [choice axiom]
25. ((v1_xboole_0(sK4) & ~v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12)))) | (v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12))) & ~v1_xboole_0(sK4)) | (v1_xboole_0(sK6) & ~v1_xboole_0(k2_gate_1(sK3,sK12))) | (v1_xboole_0(k2_gate_1(sK3,sK12)) & ~v1_xboole_0(sK6)) | (v1_xboole_0(sK7) & ~v1_xboole_0(k2_gate_1(sK1,sK12))) | (v1_xboole_0(k2_gate_1(sK1,sK12)) & ~v1_xboole_0(sK7)) | (v1_xboole_0(sK5) & ~v1_xboole_0(k2_gate_1(sK0,sK12))) | (v1_xboole_0(k2_gate_1(sK0,sK12)) & ~v1_xboole_0(sK5))) & (~v1_xboole_0(sK11) | v1_xboole_0(k2_gate_1(sK8,sK12))) & (~v1_xboole_0(k2_gate_1(sK8,sK12)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK12))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK12)) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK7) | v1_xboole_0(k2_gate_1(sK11,sK10))) & (~v1_xboole_0(k2_gate_1(sK11,sK10)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10)))) & (~v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k2_gate_1(sK9,sK8))) & (~v1_xboole_0(k2_gate_1(sK9,sK8)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8)))) & (~v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8)))) & (~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) | v1_xboole_0(sK0)) [skolemisation 15,24]
26. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) | v1_xboole_0(sK0) [cnf transformation 25] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. ~v1_xboole_0(sK0) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) [cnf transformation 25] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) | v1_xboole_0(sK1) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~v1_xboole_0(sK1) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. ~v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) | v1_xboole_0(sK2) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~v1_xboole_0(sK2) | v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. ~v1_xboole_0(k2_gate_1(sK9,sK8)) | v1_xboole_0(sK3) [cnf transformation 25] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. ~v1_xboole_0(sK3) | v1_xboole_0(k2_gate_1(sK9,sK8)) [cnf transformation 25] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) | v1_xboole_0(sK4) [cnf transformation 25] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. ~v1_xboole_0(sK4) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) [cnf transformation 25] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) | v1_xboole_0(sK5) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. ~v1_xboole_0(sK5) | v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. ~v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) | v1_xboole_0(sK6) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. ~v1_xboole_0(sK6) | v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. ~v1_xboole_0(k2_gate_1(sK11,sK10)) | v1_xboole_0(sK7) [cnf transformation 25] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
41. ~v1_xboole_0(sK7) | v1_xboole_0(k2_gate_1(sK11,sK10)) [cnf transformation 25] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK12)) | v1_xboole_0(sK10) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. ~v1_xboole_0(sK10) | v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK12)) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
44. ~v1_xboole_0(k2_gate_1(sK8,sK12)) | v1_xboole_0(sK11) [cnf transformation 25] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. ~v1_xboole_0(sK11) | v1_xboole_0(k2_gate_1(sK8,sK12)) [cnf transformation 25] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12))) | ~v1_xboole_0(sK4) | ~v1_xboole_0(k2_gate_1(sK3,sK12)) | ~v1_xboole_0(sK6) | ~v1_xboole_0(k2_gate_1(sK1,sK12)) | ~v1_xboole_0(sK7) | v1_xboole_0(sK5) | v1_xboole_0(k2_gate_1(sK0,sK12)) [cnf transformation 25] {a:0,w:27,wCS:2700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12))) | ~v1_xboole_0(sK4) | ~v1_xboole_0(k2_gate_1(sK3,sK12)) | ~v1_xboole_0(sK6) | v1_xboole_0(sK7) | v1_xboole_0(k2_gate_1(sK1,sK12)) | ~v1_xboole_0(k2_gate_1(sK0,sK12)) | ~v1_xboole_0(sK5) [cnf transformation 25] {a:0,w:27,wCS:2700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
94. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12))) | ~v1_xboole_0(sK4) | v1_xboole_0(sK6) | v1_xboole_0(k2_gate_1(sK3,sK12)) | ~v1_xboole_0(k2_gate_1(sK1,sK12)) | ~v1_xboole_0(sK7) | ~v1_xboole_0(k2_gate_1(sK0,sK12)) | ~v1_xboole_0(sK5) [cnf transformation 25] {a:0,w:27,wCS:2700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
238. v1_xboole_0(sK4) | v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12))) | ~v1_xboole_0(k2_gate_1(sK3,sK12)) | ~v1_xboole_0(sK6) | ~v1_xboole_0(k2_gate_1(sK1,sK12)) | ~v1_xboole_0(sK7) | ~v1_xboole_0(k2_gate_1(sK0,sK12)) | ~v1_xboole_0(sK5) [cnf transformation 25] {a:0,w:27,wCS:2700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
302. v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0) [cnf transformation 16] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
303. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 17] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
304. k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
305. k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
306. ~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1) [cnf transformation 18] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
307. v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
308. ~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 21] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
309. v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 23] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
311. 1 <=> v1_xboole_0(k2_gate_1(sK0,sK12)) [avatar definition]
312. ~v1_xboole_0(k2_gate_1(sK0,sK12)) <- (~1) [avatar component clause 311] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
313. v1_xboole_0(k2_gate_1(sK0,sK12)) <- (1) [avatar component clause 311] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
315. 2 <=> v1_xboole_0(sK5) [avatar definition]
319. 3 <=> v1_xboole_0(k2_gate_1(sK1,sK12)) [avatar definition]
320. ~v1_xboole_0(k2_gate_1(sK1,sK12)) <- (~3) [avatar component clause 319] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
321. v1_xboole_0(k2_gate_1(sK1,sK12)) <- (3) [avatar component clause 319] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
323. 4 <=> v1_xboole_0(sK7) [avatar definition]
327. 5 <=> v1_xboole_0(k2_gate_1(sK3,sK12)) [avatar definition]
328. ~v1_xboole_0(k2_gate_1(sK3,sK12)) <- (~5) [avatar component clause 327] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
329. v1_xboole_0(k2_gate_1(sK3,sK12)) <- (5) [avatar component clause 327] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
331. 6 <=> v1_xboole_0(sK6) [avatar definition]
335. 7 <=> v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12))) [avatar definition]
336. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12))) <- (~7) [avatar component clause 335] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
337. v1_xboole_0(k3_gate_1(k2_gate_1(sK2,sK12),k1_gate_1(sK12))) <- (7) [avatar component clause 335] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
339. 8 <=> v1_xboole_0(sK4) [avatar definition]
349. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | 7 | 8 [avatar split clause 238,339,335,327,331,319,323,311,315]
353. ~2 | ~1 | ~4 | ~3 | 5 | 6 | ~8 | ~7 [avatar split clause 94,335,339,331,327,319,323,311,315]
355. ~2 | ~1 | 3 | 4 | ~6 | ~5 | ~8 | ~7 [avatar split clause 58,335,339,327,331,323,319,311,315]
356. 1 | 2 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 [avatar split clause 49,335,339,327,331,319,323,315,311]
359. 9 <=> v1_xboole_0(k2_gate_1(sK8,sK12)) [avatar definition]
360. ~v1_xboole_0(k2_gate_1(sK8,sK12)) <- (~9) [avatar component clause 359] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
361. v1_xboole_0(k2_gate_1(sK8,sK12)) <- (9) [avatar component clause 359] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
363. 10 <=> v1_xboole_0(sK11) [avatar definition]
366. 9 | ~10 [avatar split clause 45,363,359]
367. 10 | ~9 [avatar split clause 44,359,363]
369. 11 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK12)) [avatar definition]
370. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK12)) <- (~11) [avatar component clause 369] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
371. v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK12)) <- (11) [avatar component clause 369] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
373. 12 <=> v1_xboole_0(sK10) [avatar definition]
374. v1_xboole_0(sK10) <- (12) [avatar component clause 373] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
375. ~v1_xboole_0(sK10) <- (~12) [avatar component clause 373] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
376. 11 | ~12 [avatar split clause 43,373,369]
377. 12 | ~11 [avatar split clause 42,369,373]
379. 13 <=> v1_xboole_0(k2_gate_1(sK11,sK10)) [avatar definition]
380. ~v1_xboole_0(k2_gate_1(sK11,sK10)) <- (~13) [avatar component clause 379] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
381. v1_xboole_0(k2_gate_1(sK11,sK10)) <- (13) [avatar component clause 379] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
382. 13 | ~4 [avatar split clause 41,323,379]
383. 4 | ~13 [avatar split clause 40,379,323]
385. 14 <=> v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) [avatar definition]
386. ~v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) <- (~14) [avatar component clause 385] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
387. v1_xboole_0(k2_gate_1(sK11,k1_gate_1(sK10))) <- (14) [avatar component clause 385] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
388. 14 | ~6 [avatar split clause 39,331,385]
389. 6 | ~14 [avatar split clause 38,385,331]
391. 15 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) [avatar definition]
392. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) <- (~15) [avatar component clause 391] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
393. v1_xboole_0(k2_gate_1(k1_gate_1(sK11),sK10)) <- (15) [avatar component clause 391] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
394. 15 | ~2 [avatar split clause 37,315,391]
395. 2 | ~15 [avatar split clause 36,391,315]
397. 16 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) [avatar definition]
398. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) <- (~16) [avatar component clause 397] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
399. v1_xboole_0(k2_gate_1(k1_gate_1(sK11),k1_gate_1(sK10))) <- (16) [avatar component clause 397] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
400. 16 | ~8 [avatar split clause 35,339,397]
401. 8 | ~16 [avatar split clause 34,397,339]
403. 17 <=> v1_xboole_0(k2_gate_1(sK9,sK8)) [avatar definition]
404. ~v1_xboole_0(k2_gate_1(sK9,sK8)) <- (~17) [avatar component clause 403] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
405. v1_xboole_0(k2_gate_1(sK9,sK8)) <- (17) [avatar component clause 403] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
407. 18 <=> v1_xboole_0(sK3) [avatar definition]
408. v1_xboole_0(sK3) <- (18) [avatar component clause 407] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
409. ~v1_xboole_0(sK3) <- (~18) [avatar component clause 407] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
410. 17 | ~18 [avatar split clause 33,407,403]
411. 18 | ~17 [avatar split clause 32,403,407]
413. 19 <=> v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) [avatar definition]
414. ~v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) <- (~19) [avatar component clause 413] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
415. v1_xboole_0(k2_gate_1(sK9,k1_gate_1(sK8))) <- (19) [avatar component clause 413] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
417. 20 <=> v1_xboole_0(sK2) [avatar definition]
418. v1_xboole_0(sK2) <- (20) [avatar component clause 417] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
420. 19 | ~20 [avatar split clause 31,417,413]
421. 20 | ~19 [avatar split clause 30,413,417]
423. 21 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) [avatar definition]
424. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) <- (~21) [avatar component clause 423] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
425. v1_xboole_0(k2_gate_1(k1_gate_1(sK9),sK8)) <- (21) [avatar component clause 423] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
427. 22 <=> v1_xboole_0(sK1) [avatar definition]
428. v1_xboole_0(sK1) <- (22) [avatar component clause 427] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
430. 21 | ~22 [avatar split clause 29,427,423]
431. 22 | ~21 [avatar split clause 28,423,427]
433. 23 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) [avatar definition]
434. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) <- (~23) [avatar component clause 433] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
435. v1_xboole_0(k2_gate_1(k1_gate_1(sK9),k1_gate_1(sK8))) <- (23) [avatar component clause 433] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
437. 24 <=> v1_xboole_0(sK0) [avatar definition]
439. ~v1_xboole_0(sK0) <- (~24) [avatar component clause 437] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
440. 23 | ~24 [avatar split clause 27,437,433]
441. 24 | ~23 [avatar split clause 26,433,437]
443. ~v1_xboole_0(k1_gate_1(sK11)) <- (~16) [resolution 307,398] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
444. ~v1_xboole_0(k1_gate_1(sK9)) <- (~23) [resolution 307,434] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
445. ~v1_xboole_0(k1_gate_1(sK9)) <- (~21) [resolution 307,424] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
446. ~v1_xboole_0(k1_gate_1(sK11)) <- (~15) [resolution 307,392] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
447. ~v1_xboole_0(k1_gate_1(sK9)) <- (~11) [resolution 307,370] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
448. ~v1_xboole_0(sK8) <- (~9) [resolution 307,360] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
449. ~v1_xboole_0(sK9) <- (~17) [resolution 307,404] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
451. ~v1_xboole_0(sK11) <- (~13) [resolution 307,380] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
452. ~v1_xboole_0(sK11) <- (~14) [resolution 307,386] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
454. 25 <=> v1_xboole_0(k1_gate_1(sK11)) [avatar definition]
455. v1_xboole_0(k1_gate_1(sK11)) <- (25) [avatar component clause 454] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
456. ~v1_xboole_0(k1_gate_1(sK11)) <- (~25) [avatar component clause 454] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
457. ~25 | 16 [avatar split clause 443,397,454]
459. 26 <=> v1_xboole_0(k1_gate_1(sK9)) [avatar definition]
460. v1_xboole_0(k1_gate_1(sK9)) <- (26) [avatar component clause 459] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
461. ~v1_xboole_0(k1_gate_1(sK9)) <- (~26) [avatar component clause 459] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
462. ~26 | 23 [avatar split clause 444,433,459]
463. ~26 | 21 [avatar split clause 445,423,459]
464. ~25 | 15 [avatar split clause 446,391,454]
465. ~26 | 11 [avatar split clause 447,369,459]
467. 27 <=> v1_xboole_0(sK8) [avatar definition]
468. v1_xboole_0(sK8) <- (27) [avatar component clause 467] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
469. ~v1_xboole_0(sK8) <- (~27) [avatar component clause 467] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
470. ~27 | 9 [avatar split clause 448,359,467]
472. 28 <=> v1_xboole_0(sK9) [avatar definition]
474. ~v1_xboole_0(sK9) <- (~28) [avatar component clause 472] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
475. ~28 | 17 [avatar split clause 449,403,472]
477. ~v1_xboole_0(k3_gate_1(X1,X0)) | v1_xboole_0(X1) [superposition 306,304] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
486. v1_xboole_0(k2_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 307,305] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
489. 29 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) [avatar definition]
490. v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) <- (29) [avatar component clause 489] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
491. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) <- (~29) [avatar component clause 489] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
494. 30 <=> v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) [avatar definition]
495. v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) <- (30) [avatar component clause 494] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
496. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) <- (~30) [avatar component clause 494] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
499. 31 <=> v1_xboole_0(k2_gate_1(sK8,sK9)) [avatar definition]
500. v1_xboole_0(k2_gate_1(sK8,sK9)) <- (31) [avatar component clause 499] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
501. ~v1_xboole_0(k2_gate_1(sK8,sK9)) <- (~31) [avatar component clause 499] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
504. 32 <=> v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) [avatar definition]
505. v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) <- (32) [avatar component clause 504] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
506. ~v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) <- (~32) [avatar component clause 504] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
509. 33 <=> v1_xboole_0(k2_gate_1(sK10,sK11)) [avatar definition]
510. v1_xboole_0(k2_gate_1(sK10,sK11)) <- (33) [avatar component clause 509] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
511. ~v1_xboole_0(k2_gate_1(sK10,sK11)) <- (~33) [avatar component clause 509] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
514. 34 <=> v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) [avatar definition]
515. v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) <- (34) [avatar component clause 514] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
516. ~v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) <- (~34) [avatar component clause 514] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
519. 35 <=> v1_xboole_0(k2_gate_1(sK12,k1_gate_1(sK9))) [avatar definition]
520. v1_xboole_0(k2_gate_1(sK12,k1_gate_1(sK9))) <- (35) [avatar component clause 519] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
521. ~v1_xboole_0(k2_gate_1(sK12,k1_gate_1(sK9))) <- (~35) [avatar component clause 519] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
523. ~v1_xboole_0(sK8) <- (~31) [resolution 501,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
526. v1_xboole_0(sK11) <- (~25) [resolution 456,302] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
530. 10 | 25 [avatar split clause 526,454,363]
531. ~10 | 14 [avatar split clause 452,385,363]
532. ~10 | 13 [avatar split clause 451,379,363]
539. v1_xboole_0(sK9) <- (~26) [resolution 461,302] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
541. $false <- (~26, ~28) [subsumption resolution 539,474] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
542. 26 | 28 [avatar contradiction clause 541]
550. ~v1_xboole_0(sK10) <- (~34) [resolution 516,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
552. $false <- (12, ~34) [subsumption resolution 550,374] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
553. ~12 | 34 [avatar contradiction clause 552]
556. ~v1_xboole_0(k1_gate_1(sK10)) <- (~29) [resolution 491,307] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
558. 36 <=> v1_xboole_0(k1_gate_1(sK10)) [avatar definition]
559. v1_xboole_0(k1_gate_1(sK10)) <- (36) [avatar component clause 558] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
560. ~v1_xboole_0(k1_gate_1(sK10)) <- (~36) [avatar component clause 558] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
561. ~36 | 29 [avatar split clause 556,489,558]
564. v1_xboole_0(sK12) | v1_xboole_0(sK1) <- (3) [resolution 308,321] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
566. v1_xboole_0(k1_gate_1(sK9)) | v1_xboole_0(sK8) <- (32) [resolution 308,505] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
568. v1_xboole_0(k1_gate_1(sK11)) | v1_xboole_0(sK10) <- (34) [resolution 308,515] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
569. v1_xboole_0(k1_gate_1(sK10)) | v1_xboole_0(sK11) <- (14) [resolution 308,387] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
570. v1_xboole_0(k1_gate_1(sK9)) | v1_xboole_0(sK12) <- (35) [resolution 308,520] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
574. 37 <=> v1_xboole_0(k1_gate_1(sK8)) [avatar definition]
575. ~v1_xboole_0(k1_gate_1(sK8)) <- (~37) [avatar component clause 574] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
576. v1_xboole_0(k1_gate_1(sK8)) <- (37) [avatar component clause 574] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
580. 38 <=> v1_xboole_0(sK12) [avatar definition]
581. ~v1_xboole_0(sK12) <- (~38) [avatar component clause 580] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
582. v1_xboole_0(sK12) <- (38) [avatar component clause 580] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
586. 38 | 26 | ~35 [avatar split clause 570,519,459,580]
587. 10 | 36 | ~14 [avatar split clause 569,385,558,363]
594. 39 <=> v1_xboole_0(k3_gate_1(k1_gate_1(sK12),k2_gate_1(sK2,sK12))) [avatar definition]
595. v1_xboole_0(k3_gate_1(k1_gate_1(sK12),k2_gate_1(sK2,sK12))) <- (39) [avatar component clause 594] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
596. ~v1_xboole_0(k3_gate_1(k1_gate_1(sK12),k2_gate_1(sK2,sK12))) <- (~39) [avatar component clause 594] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
600. ~v1_xboole_0(sK11) <- (25) [resolution 455,303] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
601. v1_xboole_0(k1_gate_1(sK11)) | v1_xboole_0(k1_gate_1(sK10)) <- (29) [resolution 490,308] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
604. ~v1_xboole_0(k2_gate_1(sK2,sK12)) | ~v1_xboole_0(k1_gate_1(sK12)) <- (~39) [resolution 309,596] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
608. 40 <=> v1_xboole_0(k1_gate_1(sK12)) [avatar definition]
609. v1_xboole_0(k1_gate_1(sK12)) <- (40) [avatar component clause 608] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
610. ~v1_xboole_0(k1_gate_1(sK12)) <- (~40) [avatar component clause 608] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
612. 41 <=> v1_xboole_0(k2_gate_1(sK2,sK12)) [avatar definition]
613. v1_xboole_0(k2_gate_1(sK2,sK12)) <- (41) [avatar component clause 612] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
614. ~v1_xboole_0(k2_gate_1(sK2,sK12)) <- (~41) [avatar component clause 612] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
615. ~40 | ~41 | 39 [avatar split clause 604,594,612,608]
617. v1_xboole_0(sK12) <- (~40) [resolution 610,302] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
618. 38 | 40 [avatar split clause 617,608,580]
625. ~v1_xboole_0(sK12) <- (~5) [resolution 486,328] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
634. $false <- (~5, 38) [subsumption resolution 625,582] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
635. 5 | ~38 [avatar contradiction clause 634]
638. ~10 | ~25 [avatar split clause 600,454,363]
642. 36 | 25 | ~29 [avatar split clause 601,489,454,558]
647. ~v1_xboole_0(sK10) <- (36) [resolution 559,303] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
649. $false <- (12, 36) [subsumption resolution 647,374] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
650. ~12 | ~36 [avatar contradiction clause 649]
651. ~v1_xboole_0(sK12) <- (~9) [resolution 360,486] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
653. ~38 | 9 [avatar split clause 651,359,580]
654. ~v1_xboole_0(k1_gate_1(sK10)) <- (~14) [resolution 386,486] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
656. ~36 | 14 [avatar split clause 654,385,558]
660. v1_xboole_0(sK8) <- (~37) [resolution 575,302] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
662. $false <- (~27, ~37) [subsumption resolution 660,469] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
663. 27 | 37 [avatar contradiction clause 662]
664. v1_xboole_0(k1_gate_1(sK8)) | v1_xboole_0(sK9) <- (19) [resolution 415,308] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
665. v1_xboole_0(k1_gate_1(sK8)) <- (19, ~28) [subsumption resolution 664,474] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
666. 37 | ~19 | 28 [avatar split clause 665,472,413,574]
667. ~v1_xboole_0(sK8) <- (37) [resolution 576,303] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
671. ~v1_xboole_0(sK2) <- (~41) [resolution 614,307] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
673. $false <- (20, ~41) [subsumption resolution 671,418] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
674. ~20 | 41 [avatar contradiction clause 673]
678. ~v1_xboole_0(sK12) <- (~1) [resolution 312,486] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
679. ~v1_xboole_0(sK0) <- (~1) [resolution 312,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
683. ~24 | 1 [avatar split clause 679,311,437]
684. ~27 | ~37 [avatar split clause 667,574,467]
685. 27 | 26 | ~32 [avatar split clause 566,504,459,467]
686. ~27 | 31 [avatar split clause 523,499,467]
693. 22 | 38 | ~3 [avatar split clause 564,319,580,427]
697. ~38 | 1 [avatar split clause 678,311,580]
702. v1_xboole_0(sK12) | v1_xboole_0(sK0) <- (1) [resolution 313,308] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
703. v1_xboole_0(sK12) <- (1, ~24) [subsumption resolution 702,439] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
704. 38 | ~1 | 24 [avatar split clause 703,437,311,580]
705. v1_xboole_0(sK12) | v1_xboole_0(sK8) <- (9) [resolution 361,308] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
708. v1_xboole_0(sK9) | v1_xboole_0(sK8) <- (31) [resolution 500,308] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
721. ~v1_xboole_0(k1_gate_1(sK8)) <- (~30) [resolution 496,307] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
723. $false <- (~30, 37) [subsumption resolution 721,576] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
724. 30 | ~37 [avatar contradiction clause 723]
725. v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) <- (23) [forward demodulation 435,305] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
726. 30 | ~23 [avatar split clause 725,433,494]
728. v1_xboole_0(k1_gate_1(sK9)) | v1_xboole_0(k1_gate_1(sK8)) <- (30) [resolution 495,308] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
735. v1_xboole_0(sK10) <- (~25, 34) [subsumption resolution 568,456] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
736. 12 | 25 | ~34 [avatar split clause 735,514,454,373]
741. ~v1_xboole_0(k3_gate_1(k1_gate_1(sK12),k2_gate_1(sK2,sK12))) <- (~7) [forward demodulation 336,304] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
742. ~39 | 7 [avatar split clause 741,335,594]
751. ~v1_xboole_0(sK12) <- (~3) [resolution 320,486] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
752. ~v1_xboole_0(sK1) <- (~3) [resolution 320,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
755. $false <- (~3, 38) [subsumption resolution 751,582] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
756. 3 | ~38 [avatar contradiction clause 755]
757. $false <- (~3, 22) [subsumption resolution 752,428] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
758. 3 | ~22 [avatar contradiction clause 757]
769. v1_xboole_0(sK10) <- (~36) [resolution 560,302] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
771. $false <- (~12, ~36) [subsumption resolution 769,375] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
772. 12 | 36 [avatar contradiction clause 771]
780. ~v1_xboole_0(sK12) <- (~35) [resolution 521,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
782. $false <- (~35, 38) [subsumption resolution 780,582] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
783. 35 | ~38 [avatar contradiction clause 782]
784. 27 | 28 | ~31 [avatar split clause 708,499,472,467]
785. 27 | 38 | ~9 [avatar split clause 705,359,580,467]
786. 37 | 26 | ~30 [avatar split clause 728,494,459,574]
787. v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) <- (21) [forward demodulation 425,305] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
788. 32 | ~21 [avatar split clause 787,423,504]
789. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) <- (~16) [forward demodulation 398,305] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
790. ~29 | 16 [avatar split clause 789,397,489]
791. v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) <- (15) [forward demodulation 393,305] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
792. 34 | ~15 [avatar split clause 791,391,514]
793. v1_xboole_0(k2_gate_1(sK12,k1_gate_1(sK9))) <- (11) [forward demodulation 371,305] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
794. 35 | ~11 [avatar split clause 793,369,519]
797. ~v1_xboole_0(sK3) <- (~5) [resolution 328,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
799. $false <- (~5, 18) [subsumption resolution 797,408] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
800. 5 | ~18 [avatar contradiction clause 799]
801. v1_xboole_0(k3_gate_1(k1_gate_1(sK12),k2_gate_1(sK2,sK12))) <- (7) [forward demodulation 337,304] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
802. 39 | ~7 [avatar split clause 801,335,594]
803. ~v1_xboole_0(k1_gate_1(sK8)) <- (~19) [resolution 414,486] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
804. ~v1_xboole_0(sK9) <- (~19) [resolution 414,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
805. ~37 | 19 [avatar split clause 803,413,574]
806. ~28 | 19 [avatar split clause 804,413,472]
808. ~v1_xboole_0(sK9) <- (26) [resolution 460,303] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
815. v1_xboole_0(k1_gate_1(sK12)) <- (39) [resolution 595,477] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
816. v1_xboole_0(k2_gate_1(sK2,sK12)) <- (39) [resolution 595,306] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
817. 40 | ~39 [avatar split clause 815,594,608]
818. 41 | ~39 [avatar split clause 816,594,612]
819. ~v1_xboole_0(sK12) <- (40) [resolution 609,303] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
820. ~38 | ~40 [avatar split clause 819,608,580]
822. v1_xboole_0(sK12) | v1_xboole_0(sK2) <- (41) [resolution 613,308] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
824. v1_xboole_0(sK2) <- (~38, 41) [subsumption resolution 822,581] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
827. 20 | 38 | ~41 [avatar split clause 824,612,580,417]
828. ~28 | ~26 [avatar split clause 808,459,472]
829. ~v1_xboole_0(k2_gate_1(sK8,sK9)) <- (~17) [forward demodulation 404,305] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
830. ~31 | 17 [avatar split clause 829,403,499]
831. v1_xboole_0(k2_gate_1(k1_gate_1(sK10),k1_gate_1(sK11))) <- (16) [forward demodulation 399,305] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
832. 29 | ~16 [avatar split clause 831,397,489]
834. v1_xboole_0(sK12) | v1_xboole_0(sK3) <- (5) [resolution 329,308] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
836. v1_xboole_0(sK3) <- (5, ~38) [subsumption resolution 834,581] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
837. $false <- (5, ~18, ~38) [subsumption resolution 836,409] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
838. ~5 | 18 | 38 [avatar contradiction clause 837]
839. ~v1_xboole_0(k2_gate_1(sK10,sK11)) <- (~13) [forward demodulation 380,305] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
840. ~33 | 13 [avatar split clause 839,379,509]
850. ~v1_xboole_0(sK10) <- (~33) [resolution 511,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
852. $false <- (12, ~33) [subsumption resolution 850,374] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
853. ~12 | 33 [avatar contradiction clause 852]
854. ~v1_xboole_0(k2_gate_1(sK8,k1_gate_1(sK9))) <- (~21) [forward demodulation 424,305] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
855. ~32 | 21 [avatar split clause 854,423,504]
856. v1_xboole_0(k2_gate_1(sK8,sK9)) <- (17) [forward demodulation 405,305] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
857. 31 | ~17 [avatar split clause 856,403,499]
858. v1_xboole_0(k2_gate_1(sK10,sK11)) <- (13) [forward demodulation 381,305] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
859. 33 | ~13 [avatar split clause 858,379,509]
860. ~v1_xboole_0(k2_gate_1(sK12,k1_gate_1(sK9))) <- (~11) [forward demodulation 370,305] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
861. ~35 | 11 [avatar split clause 860,369,519]
867. v1_xboole_0(sK11) | v1_xboole_0(sK10) <- (33) [resolution 510,308] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
872. v1_xboole_0(sK11) <- (~12, 33) [subsumption resolution 867,375] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
873. 10 | 12 | ~33 [avatar split clause 872,509,373,363]
874. ~v1_xboole_0(k2_gate_1(k1_gate_1(sK8),k1_gate_1(sK9))) <- (~23) [forward demodulation 434,305] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
875. ~30 | 23 [avatar split clause 874,433,494]
876. ~v1_xboole_0(k2_gate_1(sK10,k1_gate_1(sK11))) <- (~15) [forward demodulation 392,305] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
877. ~34 | 15 [avatar split clause 876,391,514]
888. ~v1_xboole_0(sK8) <- (~32) [resolution 506,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
890. $false <- (27, ~32) [subsumption resolution 888,468] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
891. ~27 | 32 [avatar contradiction clause 890]
892. $false [avatar sat refutation 349,353,355,356,366,367,376,377,382,383,388,389,394,395,400,401,410,411,420,421,430,431,440,441,457,462,463,464,465,470,475,530,531,532,542,553,561,586,587,615,618,635,638,642,650,653,656,663,666,674,683,684,685,686,693,697,704,724,726,736,742,756,758,772,783,784,785,786,788,790,792,794,800,802,805,806,817,818,820,827,828,830,832,838,840,853,855,857,859,861,873,875,877,891] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:911,thDist:-911}
% SZS output end Proof for gate_3__t2_gate_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 13
% 
% >>> Saturation
% Initial clauses: 284
% Generated clauses: 790
% Active clauses: 179
% Passive clauses: 203
% Final active clauses: 38
% Final passive clauses: 5
% Main loop iterations started: 180
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 37
% Fw demodulations: 37
% Bw demodulations: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 245
% Forward subsumptions: 89
% 
% >>> Generating Inferences
% Binary resolution: 129
% Unit resulting resolution: 21
% Forward superposition: 8
% Backward superposition: 4
% 
% >>> AVATAR
% Split clauses: 46
% Split components: 188
% Sat splitting refutations: 46
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 177
% SAT solver binary clauses: 133
% 
% Memory used [KB]: 9978
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s
% superposition: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

