i: [1,46,0,13,3,1,0,0] t64_zfmisc_1
a: [3,72,0,3,1,46]
s: 72
a: [3,76,0,3,1,46]
s: 76
a: [3,69,0,7,1,46]
s: 69
i: [1,47,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,84,0,2,1,47]
s: 84
i: [1,48,0,2,1,0,0,1] fc1_xboole_0
a: [3,89,0,2,1,48]
s: 89
i: [1,64,0,2,1,0,0,2] rc2_xboole_0
a: [3,99,0,2,1,64]
s: 99
i: [1,65,0,2,1,0,0,2] rc1_xboole_0
a: [3,104,0,2,1,65]
s: 104
i: [1,60,0,3,1,0,0,255] reflexivity_r1_tarski
s: 60
i: [1,49,0,3,1,0,0,255] d2_xboole_0
a: [3,94,0,3,1,49]
s: 94
r: 84
i: [1,61,0,4,1,0,0,1] fc3_xboole_0
s: 61
i: [1,50,0,5,1,0,0,1] t4_boole
s: 50
i: [1,51,0,5,1,0,0,1] t3_boole
s: 51
i: [1,52,0,5,2,0,0,1] t6_boole
s: 52
i: [1,56,0,5,2,0,0,1] t7_boole
s: 56
i: [1,63,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 63
i: [1,55,0,7,3,0,0,1] t8_boole
s: 55
d: [2,111,1,2,1,1,40,56,72]
a: [3,116,1,2,1,111]
s: 116
i: [1,62,0,7,1,0,0,1] commutativity_k2_tarski
s: 62
d: [2,119,1,3,1,1,40,63,72]
a: [3,124,1,3,1,119]
s: 124
i: [1,53,0,8,2,0,0,255] t37_xboole_1
s: 53
i: [1,54,0,8,2,0,0,255] t37_xboole_1
s: 54
d: [2,120,1,3,1,1,40,63,76]
a: [3,129,1,3,1,120]
s: 129
i: [1,57,0,8,2,0,0,255] t32_zfmisc_1
s: 57
d: [2,137,1,8,2,1,44,69,54]
d: [2,140,1,5,1,1,19,137]
a: [3,144,1,5,1,140]
s: 144
i: [1,58,0,8,2,0,0,255] t32_zfmisc_1
s: 58
i: [1,59,0,11,3,0,0,255] t32_zfmisc_1
s: 59
d: [2,152,2,0,0,3,52,72,76,144,59]
e: 152
d: [2,155,2,6,2,1,40,59,144]
d: [2,159,2,3,1,2,21,155,76]
d: [2,160,2,0,0,3,21,159,72]
e: 160
r: 69
r: 76
a: [3,77,0,3,1,46]
s: 77
a: [3,68,0,7,1,46]
s: 68
d: [2,164,2,0,0,2,52,144,68,53]
e: 164
d: [2,165,1,8,2,1,44,53,68]
d: [2,166,1,5,1,1,19,165]
d: [2,168,1,0,0,2,21,166,144]
e: 168
r: 144
a: [3,143,1,5,1,140]
s: 143
d: [2,171,2,0,0,2,52,77,143,58]
e: 171
d: [2,172,2,3,1,1,40,143,58]
d: [2,175,2,0,0,2,21,172,77]
e: 175
r: 72
r: 77
i: [1,45,0,10,2,1,0,0] t64_zfmisc_1
i: [1,44,0,10,2,1,0,0] t64_zfmisc_1
d: [2,162,2,3,1,2,21,155,72]
d: [2,173,2,3,1,1,40,143,57]
f: 172,173,166,162,140,44,45,46
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t64_zfmisc_1
% SZS output start Proof for zfmisc_1__t64_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (k4_xboole_0(k2_tarski(X0,X1),X2) = k1_xboole_0 <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [input t64_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (k4_xboole_0(k2_tarski(X0,X1),X2) = k1_xboole_0 <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [negated conjecture 1]
15. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),X2) <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [input t32_zfmisc_1]
16. ! [X0] : ! [X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [input t37_xboole_1]
22. ~! [X0,X1,X2] : (k4_xboole_0(k2_tarski(X0,X1),X2) = k1_xboole_0 <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [flattening 2]
23. ! [X0,X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [flattening 16]
26. ! [X0,X1,X2] : (r1_tarski(k2_tarski(X0,X1),X2) <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [flattening 15]
28. ? [X0,X1,X2] : (k4_xboole_0(k2_tarski(X0,X1),X2) = k1_xboole_0 <~> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [ennf transformation 22]
33. ? [X0,X1,X2] : (((~r2_hidden(X1,X2) | ~r2_hidden(X0,X2)) | k4_xboole_0(k2_tarski(X0,X1),X2) != k1_xboole_0) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | k4_xboole_0(k2_tarski(X0,X1),X2) = k1_xboole_0)) [nnf transformation 28]
34. ? [X0,X1,X2] : ((~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | k4_xboole_0(k2_tarski(X0,X1),X2) != k1_xboole_0) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | k4_xboole_0(k2_tarski(X0,X1),X2) = k1_xboole_0)) [flattening 33]
35. ? [X0,X1,X2] : ((~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | k4_xboole_0(k2_tarski(X0,X1),X2) != k1_xboole_0) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | k4_xboole_0(k2_tarski(X0,X1),X2) = k1_xboole_0)) => ((~r2_hidden(sK1,sK2) | ~r2_hidden(sK0,sK2) | k1_xboole_0 != k4_xboole_0(k2_tarski(sK0,sK1),sK2)) & ((r2_hidden(sK1,sK2) & r2_hidden(sK0,sK2)) | k1_xboole_0 = k4_xboole_0(k2_tarski(sK0,sK1),sK2))) [choice axiom]
36. (~r2_hidden(sK1,sK2) | ~r2_hidden(sK0,sK2) | k1_xboole_0 != k4_xboole_0(k2_tarski(sK0,sK1),sK2)) & ((r2_hidden(sK1,sK2) & r2_hidden(sK0,sK2)) | k1_xboole_0 = k4_xboole_0(k2_tarski(sK0,sK1),sK2)) [skolemisation 34,35]
37. ! [X0,X1] : ((k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | k1_xboole_0 != k4_xboole_0(X0,X1))) [nnf transformation 23]
38. ! [X0,X1,X2] : ((r1_tarski(k2_tarski(X0,X1),X2) | (~r2_hidden(X1,X2) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | ~r1_tarski(k2_tarski(X0,X1),X2))) [nnf transformation 26]
39. ! [X0,X1,X2] : ((r1_tarski(k2_tarski(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | ~r1_tarski(k2_tarski(X0,X1),X2))) [flattening 38]
44. r2_hidden(sK0,sK2) | k1_xboole_0 = k4_xboole_0(k2_tarski(sK0,sK1),sK2) [cnf transformation 36] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. r2_hidden(sK1,sK2) | k1_xboole_0 = k4_xboole_0(k2_tarski(sK0,sK1),sK2) [cnf transformation 36] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. ~r2_hidden(sK1,sK2) | ~r2_hidden(sK0,sK2) | k1_xboole_0 != k4_xboole_0(k2_tarski(sK0,sK1),sK2) [cnf transformation 36] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. k1_xboole_0 != k4_xboole_0(X0,X1) | r1_tarski(X0,X1) [cnf transformation 37] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
54. k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1) [cnf transformation 37] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
57. ~r1_tarski(k2_tarski(X0,X1),X2) | r2_hidden(X0,X2) [cnf transformation 39] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
58. ~r1_tarski(k2_tarski(X0,X1),X2) | r2_hidden(X1,X2) [cnf transformation 39] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
59. r1_tarski(k2_tarski(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) [cnf transformation 39] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
67. 1 <=> k1_xboole_0 = k4_xboole_0(k2_tarski(sK0,sK1),sK2) [avatar definition]
68. k1_xboole_0 = k4_xboole_0(k2_tarski(sK0,sK1),sK2) <- (1) [avatar component clause 67] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
69. k1_xboole_0 != k4_xboole_0(k2_tarski(sK0,sK1),sK2) <- (~1) [avatar component clause 67] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. 2 <=> r2_hidden(sK0,sK2) [avatar definition]
72. r2_hidden(sK0,sK2) <- (2) [avatar component clause 71] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. 3 <=> r2_hidden(sK1,sK2) [avatar definition]
78. ~1 | ~2 | ~3 [avatar split clause 46,75,71,67]
79. 1 | 3 [avatar split clause 45,75,67]
80. 1 | 2 [avatar split clause 44,71,67]
137. k1_xboole_0 != k1_xboole_0 | ~r1_tarski(k2_tarski(sK0,sK1),sK2) <- (~1) [superposition 69,54] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
140. ~r1_tarski(k2_tarski(sK0,sK1),sK2) <- (~1) [trivial inequality removal 137] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. 12 <=> r1_tarski(k2_tarski(sK0,sK1),sK2) [avatar definition]
143. r1_tarski(k2_tarski(sK0,sK1),sK2) <- (12) [avatar component clause 142] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. ~r1_tarski(k2_tarski(sK0,sK1),sK2) <- (~12) [avatar component clause 142] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
145. ~12 | 1 [avatar split clause 140,67,142]
155. ~r2_hidden(sK1,sK2) | ~r2_hidden(sK0,sK2) <- (~12) [resolution 59,144] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
162. ~r2_hidden(sK1,sK2) <- (2, ~12) [subsumption resolution 155,72] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
163. ~3 | ~2 | 12 [avatar split clause 162,142,71,75]
165. k1_xboole_0 != k1_xboole_0 | r1_tarski(k2_tarski(sK0,sK1),sK2) <- (1) [superposition 53,68] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
166. r1_tarski(k2_tarski(sK0,sK1),sK2) <- (1) [trivial inequality removal 165] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
170. 12 | ~1 [avatar split clause 166,67,142]
172. r2_hidden(sK1,sK2) <- (12) [resolution 143,58] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
173. r2_hidden(sK0,sK2) <- (12) [resolution 143,57] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
177. 2 | ~12 [avatar split clause 173,142,71]
178. 3 | ~12 [avatar split clause 172,142,75]
179. $false [avatar sat refutation 78,79,80,145,163,170,177,178] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:105,thDist:-105}
% SZS output end Proof for zfmisc_1__t64_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 22
% Generated clauses: 85
% Active clauses: 29
% Passive clauses: 42
% Final active clauses: 23
% Final passive clauses: 13
% Main loop iterations started: 29
% 
% >>> Simplifying Inferences
% Trivial inequalities: 3
% Fw subsumption resolutions: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 3
% Forward superposition: 9
% Backward superposition: 9
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 25
% SAT solver unit clauses: 5
% SAT solver binary clauses: 14
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward demodulation index maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

