i: [1,32,0,4,2,1,0,0] t25_zfrefle1
a: [3,47,0,2,1,32]
s: 47
a: [3,50,0,2,1,32]
s: 50
i: [1,29,0,2,1,1,0,0] t25_zfrefle1
a: [3,62,0,2,1,29]
s: 62
i: [1,28,0,2,1,1,0,0] t25_zfrefle1
a: [3,67,0,2,1,28]
s: 67
i: [1,30,0,3,1,1,0,0] t25_zfrefle1
a: [3,57,0,3,1,30]
s: 57
i: [1,38,0,7,3,0,0,1] reflexivity_r2_ordinal2
a: [3,71,0,2,1,38]
s: 71
d: [2,103,1,0,0,2,40,71,67]
e: 103
d: [2,102,1,0,0,2,40,71,62]
e: 102
d: [2,101,1,0,0,2,52,62,71]
e: 101
r: 71
i: [1,39,0,2,1,0,0,1] rc1_ordinal1
a: [3,79,0,2,1,39]
s: 79
i: [1,43,0,2,1,0,0,1] rc1_ordinal2
a: [3,84,0,2,1,43]
s: 84
i: [1,42,0,2,1,0,0,1] rc1_ordinal2
a: [3,89,0,2,1,42]
s: 89
i: [1,41,0,2,1,0,0,1] rc1_ordinal2
a: [3,94,0,2,1,41]
s: 94
i: [1,40,0,2,1,0,0,1] rc1_ordinal2
a: [3,99,0,2,1,40]
s: 99
i: [1,33,0,4,2,0,0,2] cc1_ordinal1
s: 33
i: [1,34,0,4,2,0,0,2] cc1_ordinal1
s: 34
d: [2,107,1,2,1,1,40,33,67]
a: [3,114,1,2,1,107]
s: 114
a: [3,74,0,5,2,38]
s: 74
i: [1,37,0,6,3,0,0,2] cc2_ordinal1
s: 37
d: [2,108,1,2,1,1,40,33,62]
a: [3,119,1,2,1,108]
s: 119
i: [1,35,0,11,5,0,0,1] t38_ordinal4
s: 35
d: [2,126,1,2,1,1,40,34,67]
a: [3,133,1,2,1,126]
s: 133
i: [1,36,0,11,5,0,0,1] t38_ordinal4
s: 36
d: [2,150,1,0,0,5,52,67,62,47,50,57,36]
e: 150
d: [2,151,1,8,4,1,40,36,57]
d: [2,155,1,6,3,2,21,151,50]
d: [2,156,1,4,2,3,21,155,47]
d: [2,157,1,2,1,4,21,156,62]
d: [2,158,1,0,0,5,21,157,67]
e: 158
r: 47
r: 50
i: [1,31,0,4,2,1,0,0] t25_zfrefle1
d: [2,146,1,8,4,1,40,35,57]
d: [2,160,1,6,3,2,21,146,62]
d: [2,161,1,4,2,3,21,160,67]
d: [2,163,1,6,3,2,21,151,62]
d: [2,164,1,4,2,3,21,163,67]
f: 164,161,28,29,30,31,32
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfrefle1__t25_zfrefle1
% SZS output start Proof for zfrefle1__t25_zfrefle1
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_ordinal2(X0,X1) => (v4_ordinal1(X0) <=> v4_ordinal1(X1))))) [input t25_zfrefle1]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_ordinal2(X0,X1) => (v4_ordinal1(X0) <=> v4_ordinal1(X1))))) [negated conjecture 1]
8. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_ordinal2(X0,X1) => (v4_ordinal1(X0) <=> v4_ordinal1(X1))))) [input t38_ordinal4]
9. ? [X0] : (? [X1] : (((v4_ordinal1(X0) <~> v4_ordinal1(X1)) & r2_ordinal2(X0,X1)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
10. ? [X0] : (? [X1] : ((v4_ordinal1(X0) <~> v4_ordinal1(X1)) & r2_ordinal2(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 9]
12. ! [X0] : (! [X1] : (((v4_ordinal1(X0) <=> v4_ordinal1(X1)) | ~r2_ordinal2(X0,X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 8]
13. ! [X0] : (! [X1] : ((v4_ordinal1(X0) <=> v4_ordinal1(X1)) | ~r2_ordinal2(X0,X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 12]
18. ? [X0] : (? [X1] : (((~v4_ordinal1(X1) | ~v4_ordinal1(X0)) & (v4_ordinal1(X1) | v4_ordinal1(X0))) & r2_ordinal2(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [nnf transformation 10]
19. ? [X0] : (? [X1] : ((~v4_ordinal1(X1) | ~v4_ordinal1(X0)) & (v4_ordinal1(X1) | v4_ordinal1(X0)) & r2_ordinal2(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 18]
20. ? [X0] : (? [X1] : ((~v4_ordinal1(X1) | ~v4_ordinal1(X0)) & (v4_ordinal1(X1) | v4_ordinal1(X0)) & r2_ordinal2(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : ((~v4_ordinal1(X1) | ~v4_ordinal1(sK0)) & (v4_ordinal1(X1) | v4_ordinal1(sK0)) & r2_ordinal2(sK0,X1) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
21. ? [X1] : ((~v4_ordinal1(X1) | ~v4_ordinal1(sK0)) & (v4_ordinal1(X1) | v4_ordinal1(sK0)) & r2_ordinal2(sK0,X1) & v3_ordinal1(X1)) => ((~v4_ordinal1(sK1) | ~v4_ordinal1(sK0)) & (v4_ordinal1(sK1) | v4_ordinal1(sK0)) & r2_ordinal2(sK0,sK1) & v3_ordinal1(sK1)) [choice axiom]
22. ((~v4_ordinal1(sK1) | ~v4_ordinal1(sK0)) & (v4_ordinal1(sK1) | v4_ordinal1(sK0)) & r2_ordinal2(sK0,sK1) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 19,21,20]
23. ! [X0] : (! [X1] : (((v4_ordinal1(X0) | ~v4_ordinal1(X1)) & (v4_ordinal1(X1) | ~v4_ordinal1(X0))) | ~r2_ordinal2(X0,X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [nnf transformation 13]
28. v3_ordinal1(sK0) [cnf transformation 22] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
29. v3_ordinal1(sK1) [cnf transformation 22] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
30. r2_ordinal2(sK0,sK1) [cnf transformation 22] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
31. v4_ordinal1(sK1) | v4_ordinal1(sK0) [cnf transformation 22] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
32. ~v4_ordinal1(sK1) | ~v4_ordinal1(sK0) [cnf transformation 22] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
35. ~r2_ordinal2(X0,X1) | ~v4_ordinal1(X0) | v4_ordinal1(X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 23] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
36. ~r2_ordinal2(X0,X1) | ~v4_ordinal1(X1) | v4_ordinal1(X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 23] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. 1 <=> v4_ordinal1(sK0) [avatar definition]
49. 2 <=> v4_ordinal1(sK1) [avatar definition]
52. ~1 | ~2 [avatar split clause 32,49,45]
53. 1 | 2 [avatar split clause 31,49,45]
55. 3 <=> r2_ordinal2(sK0,sK1) [avatar definition]
57. r2_ordinal2(sK0,sK1) <- (3) [avatar component clause 55] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
58. 3 [avatar split clause 30,55]
60. 4 <=> v3_ordinal1(sK1) [avatar definition]
62. v3_ordinal1(sK1) <- (4) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
63. 4 [avatar split clause 29,60]
65. 5 <=> v3_ordinal1(sK0) [avatar definition]
67. v3_ordinal1(sK0) <- (5) [avatar component clause 65] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
68. 5 [avatar split clause 28,65]
146. ~v4_ordinal1(sK0) | v4_ordinal1(sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (3) [resolution 35,57] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
151. ~v4_ordinal1(sK1) | v4_ordinal1(sK0) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (3) [resolution 36,57] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
160. ~v4_ordinal1(sK0) | v4_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (3, 4) [subsumption resolution 146,62] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
161. ~v4_ordinal1(sK0) | v4_ordinal1(sK1) <- (3, 4, 5) [subsumption resolution 160,67] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
162. 2 | ~1 | ~3 | ~4 | ~5 [avatar split clause 161,65,60,55,45,49]
163. ~v4_ordinal1(sK1) | v4_ordinal1(sK0) | ~v3_ordinal1(sK0) <- (3, 4) [subsumption resolution 151,62] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
164. ~v4_ordinal1(sK1) | v4_ordinal1(sK0) <- (3, 4, 5) [subsumption resolution 163,67] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
165. 1 | ~2 | ~3 | ~4 | ~5 [avatar split clause 164,65,60,55,49,45]
166. $false [avatar sat refutation 52,53,58,63,68,162,165] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:124,thDist:-124}
% SZS output end Proof for zfrefle1__t25_zfrefle1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 7
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 16
% Generated clauses: 66
% Active clauses: 20
% Passive clauses: 25
% Final active clauses: 17
% Final passive clauses: 5
% Main loop iterations started: 20
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 8
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 16
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 8
% SAT solver binary clauses: 12
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

