i: [1,89,0,3,1,1,0,0] l23_mcart_1
a: [3,134,0,3,1,89]
s: 134
i: [1,88,0,3,1,1,0,0] l23_mcart_1
a: [3,139,0,3,1,88]
s: 139
i: [1,90,0,5,1,1,0,0] l23_mcart_1
a: [3,129,0,5,1,90]
s: 129
i: [1,92,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,144,0,2,1,92]
s: 144
i: [1,93,0,2,1,0,0,1] fc1_xboole_0
a: [3,149,0,2,1,93]
s: 149
i: [1,120,0,2,1,0,0,2] rc2_xboole_0
a: [3,159,0,2,1,120]
s: 159
i: [1,121,0,2,1,0,0,255] rc1_xtuple_0
a: [3,164,0,2,1,121]
s: 164
i: [1,122,0,2,1,0,0,2] rc1_xboole_0
a: [3,169,0,2,1,122]
s: 169
i: [1,124,0,2,1,0,0,2] rc1_relat_1
a: [3,174,0,2,1,124]
s: 174
i: [1,123,0,2,1,0,0,2] rc1_relat_1
a: [3,179,0,2,1,123]
s: 179
i: [1,125,0,2,1,0,0,255] rc2_relat_1
a: [3,184,0,2,1,125]
s: 184
i: [1,91,0,11,3,1,0,0] l23_mcart_1
s: 91
i: [1,95,0,3,1,0,0,255] fc2_xboole_0
s: 95
i: [1,94,0,3,1,0,0,255] d2_xboole_0
a: [3,154,0,3,1,94]
s: 154
r: 144
i: [1,97,0,4,2,0,0,2] cc1_relat_1
s: 97
i: [1,109,0,4,1,0,0,1] existence_m1_subset_1
s: 109
d: [2,187,1,2,1,1,40,97,169]
a: [3,192,1,2,1,187]
s: 192
i: [1,110,0,4,1,0,0,255] fc3_xboole_0
s: 110
d: [2,188,1,2,1,1,40,97,149]
a: [3,197,1,2,1,188]
s: 197
i: [1,111,0,4,1,0,0,255] fc1_xtuple_0
s: 111
i: [1,112,0,4,1,0,0,1] fc6_relat_1
s: 112
i: [1,98,0,5,2,0,0,1] t6_boole
s: 98
i: [1,106,0,5,2,0,0,2] t7_boole
s: 106
i: [1,113,0,5,1,0,0,255] fc5_relat_1
s: 113
i: [1,104,0,6,2,0,0,1] t1_subset
s: 104
i: [1,114,0,6,1,0,0,2] rd2_xtuple_0
s: 114
i: [1,115,0,6,1,0,0,2] rd1_xtuple_0
s: 115
i: [1,117,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 117
i: [1,105,0,7,3,0,0,1] t8_boole
s: 105
i: [1,116,0,7,1,0,0,255] commutativity_k2_tarski
s: 116
i: [1,103,0,8,3,0,0,1] t2_subset
s: 103
d: [2,204,1,9,2,1,40,103,129]
a: [3,213,1,4,1,204]
s: 213
i: [1,118,0,8,3,0,0,1] fc10_subset_1
s: 118
d: [2,216,2,4,2,1,40,118,213]
a: [3,220,2,2,1,216]
s: 220
d: [2,226,3,0,0,3,52,149,139,220,105]
e: 226
d: [2,227,3,0,0,3,52,149,139,220,105]
e: 227
d: [2,228,3,0,0,2,52,139,220,98]
e: 228
r: 220
i: [1,119,0,8,1,0,0,255] fc7_relat_1
s: 119
i: [1,96,0,9,2,0,0,255] rd3_xtuple_0
s: 96
a: [3,219,2,2,1,216]
s: 219
i: [1,107,0,9,2,0,0,1] t10_mcart_1
s: 107
a: [3,224,2,2,1,216]
s: 224
d: [2,244,3,0,0,3,52,149,134,224,105]
e: 244
d: [2,245,3,0,0,3,52,149,134,224,105]
e: 245
d: [2,246,3,0,0,2,52,134,224,98]
e: 246
r: 213
r: 224
i: [1,108,0,9,2,0,0,1] t10_mcart_1
s: 108
i: [1,101,0,10,1,0,0,255] d5_tarski
s: 101
a: [3,223,2,2,1,216]
s: 223
i: [1,102,0,12,3,0,0,2] t21_mcart_1
s: 102
a: [3,212,1,4,1,204]
s: 212
a: [3,209,1,5,1,204]
s: 209
d: [2,237,1,13,4,0,44,91,96]
s: 237
d: [2,286,2,10,3,0,47,237]
a: [3,290,2,2,1,286]
s: 290
d: [2,261,1,16,5,0,44,91,102]
s: 261
d: [2,266,2,4,1,1,40,209,108]
a: [3,274,2,4,1,266]
s: 274
d: [2,258,1,4,1,0,44,110,101]
s: 258
d: [2,239,1,5,2,0,44,113,96]
s: 239
d: [2,267,2,4,1,1,40,209,107]
a: [3,279,2,4,1,267]
s: 279
d: [2,205,1,6,2,0,40,103,109]
s: 205
a: [3,293,2,4,1,286]
s: 293
d: [2,199,1,7,3,0,44,98,98]
s: 199
d: [2,262,1,7,3,0,44,111,102]
s: 262
d: [2,331,3,0,0,2,52,290,112,209,262]
e: 331
d: [2,334,2,6,2,1,40,262,209]
d: [2,337,2,4,1,2,21,334,290]
d: [2,338,2,0,0,2,21,337,112]
e: 338
d: [2,306,3,4,1,1,40,274,104]
d: [2,317,3,4,1,1,40,279,104]
f: 338,317,306,286,267,266,246,228,216,204,88,89,90
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__l23_mcart_1
% SZS output start Proof for mcart_1__l23_mcart_1
1. ! [X0] : ! [X1] : ~(? [X2] : (! [X3] : (m1_subset_1(X3,X0) => ! [X4] : (m1_subset_1(X4,X1) => k4_tarski(X3,X4) != X2)) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & k1_xboole_0 != X1 & k1_xboole_0 != X0) [input l23_mcart_1]
2. ~! [X0] : ! [X1] : ~(? [X2] : (! [X3] : (m1_subset_1(X3,X0) => ! [X4] : (m1_subset_1(X4,X1) => k4_tarski(X3,X4) != X2)) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & k1_xboole_0 != X1 & k1_xboole_0 != X0) [negated conjecture 1]
20. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_zfmisc_1(X0,X1))) [input fc10_subset_1]
22. ! [X0,X1] : v1_xtuple_0(k4_tarski(X0,X1)) [input fc1_xtuple_0]
26. ! [X0,X1] : v1_relat_1(k2_zfmisc_1(X0,X1)) [input fc6_relat_1]
35. ! [X0] : (v1_xtuple_0(X0) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0) [input rd3_xtuple_0]
36. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [input t10_mcart_1]
37. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
38. ! [X0] : ! [X1] : (v1_relat_1(X1) => (r2_hidden(X0,X1) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0)) [input t21_mcart_1]
39. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
40. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
43. ~! [X0,X1] : ~(? [X2] : (! [X3] : (m1_subset_1(X3,X0) => ! [X4] : (m1_subset_1(X4,X1) => k4_tarski(X3,X4) != X2)) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & k1_xboole_0 != X1 & k1_xboole_0 != X0) [flattening 2]
45. ! [X0,X1] : (v1_relat_1(X1) => (r2_hidden(X0,X1) => k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0)) [flattening 38]
46. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 39]
47. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 37]
50. ! [X0,X1,X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [flattening 36]
54. ? [X0,X1] : (? [X2] : (! [X3] : (! [X4] : (k4_tarski(X3,X4) != X2 | ~m1_subset_1(X4,X1)) | ~m1_subset_1(X3,X0)) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & k1_xboole_0 != X1 & k1_xboole_0 != X0) [ennf transformation 43]
55. ! [X0] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0)) [ennf transformation 35]
57. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 40]
62. ! [X0,X1] : ((k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~r2_hidden(X0,X1)) | ~v1_relat_1(X1)) [ennf transformation 45]
63. ! [X0,X1] : (k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~r2_hidden(X0,X1) | ~v1_relat_1(X1)) [flattening 62]
64. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 46]
65. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 64]
66. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 47]
69. ! [X0,X1,X2] : ((r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1)) | ~r2_hidden(X0,k2_zfmisc_1(X1,X2))) [ennf transformation 50]
71. ! [X0,X1] : (~v1_xboole_0(k2_zfmisc_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 20]
72. ! [X0,X1] : (~v1_xboole_0(k2_zfmisc_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 71]
73. ? [X0,X1] : (? [X2] : (! [X3] : (! [X4] : (k4_tarski(X3,X4) != X2 | ~m1_subset_1(X4,X1)) | ~m1_subset_1(X3,X0)) & m1_subset_1(X2,k2_zfmisc_1(X0,X1))) & k1_xboole_0 != X1 & k1_xboole_0 != X0) => (? [X2] : (! [X3] : (! [X4] : (k4_tarski(X3,X4) != X2 | ~m1_subset_1(X4,sK1)) | ~m1_subset_1(X3,sK0)) & m1_subset_1(X2,k2_zfmisc_1(sK0,sK1))) & k1_xboole_0 != sK1 & k1_xboole_0 != sK0) [choice axiom]
74. ? [X2] : (! [X3] : (! [X4] : (k4_tarski(X3,X4) != X2 | ~m1_subset_1(X4,sK1)) | ~m1_subset_1(X3,sK0)) & m1_subset_1(X2,k2_zfmisc_1(sK0,sK1))) => (! [X3] : (! [X4] : (k4_tarski(X3,X4) != sK2 | ~m1_subset_1(X4,sK1)) | ~m1_subset_1(X3,sK0)) & m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1))) [choice axiom]
75. (! [X3] : (! [X4] : (k4_tarski(X3,X4) != sK2 | ~m1_subset_1(X4,sK1)) | ~m1_subset_1(X3,sK0)) & m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1))) & k1_xboole_0 != sK1 & k1_xboole_0 != sK0 [skolemisation 54,74,73]
88. k1_xboole_0 != sK0 [cnf transformation 75] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
89. k1_xboole_0 != sK1 [cnf transformation 75] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
90. m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1)) [cnf transformation 75] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
91. k4_tarski(X3,X4) != sK2 | ~m1_subset_1(X4,sK1) | ~m1_subset_1(X3,sK0) [cnf transformation 75] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
96. k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~v1_xtuple_0(X0) [cnf transformation 55] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
98. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 57] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
102. k4_tarski(k1_xtuple_0(X0),k2_xtuple_0(X0)) = X0 | ~r2_hidden(X0,X1) | ~v1_relat_1(X1) [cnf transformation 63] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
103. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 65] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 66] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
107. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k1_xtuple_0(X0),X1) [cnf transformation 69] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
108. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k2_xtuple_0(X0),X2) [cnf transformation 69] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
111. v1_xtuple_0(k4_tarski(X0,X1)) [cnf transformation 22] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
112. v1_relat_1(k2_zfmisc_1(X0,X1)) [cnf transformation 26] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
118. ~v1_xboole_0(k2_zfmisc_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 72] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
127. 1 <=> m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1)) [avatar definition]
129. m1_subset_1(sK2,k2_zfmisc_1(sK0,sK1)) <- (1) [avatar component clause 127] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
130. 1 [avatar split clause 90,127]
132. 2 <=> k1_xboole_0 = sK1 [avatar definition]
134. k1_xboole_0 != sK1 <- (~2) [avatar component clause 132] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
135. ~2 [avatar split clause 89,132]
137. 3 <=> k1_xboole_0 = sK0 [avatar definition]
139. k1_xboole_0 != sK0 <- (~3) [avatar component clause 137] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
140. ~3 [avatar split clause 88,137]
204. v1_xboole_0(k2_zfmisc_1(sK0,sK1)) | r2_hidden(sK2,k2_zfmisc_1(sK0,sK1)) <- (1) [resolution 103,129] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
207. 15 <=> r2_hidden(sK2,k2_zfmisc_1(sK0,sK1)) [avatar definition]
209. r2_hidden(sK2,k2_zfmisc_1(sK0,sK1)) <- (15) [avatar component clause 207] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
211. 16 <=> v1_xboole_0(k2_zfmisc_1(sK0,sK1)) [avatar definition]
213. v1_xboole_0(k2_zfmisc_1(sK0,sK1)) <- (16) [avatar component clause 211] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
214. 15 | 16 | ~1 [avatar split clause 204,127,211,207]
216. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (16) [resolution 118,213] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
218. 17 <=> v1_xboole_0(sK0) [avatar definition]
220. v1_xboole_0(sK0) <- (17) [avatar component clause 218] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
222. 18 <=> v1_xboole_0(sK1) [avatar definition]
224. v1_xboole_0(sK1) <- (18) [avatar component clause 222] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
225. 17 | 18 | ~16 [avatar split clause 216,211,222,218]
228. $false <- (~3, 17) [unit resulting resolution 139,220,98] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
232. 3 | ~17 [avatar contradiction clause 228]
237. sK2 != X0 | ~m1_subset_1(k2_xtuple_0(X0),sK1) | ~m1_subset_1(k1_xtuple_0(X0),sK0) | ~v1_xtuple_0(X0) [superposition 91,96] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
246. $false <- (~2, 18) [unit resulting resolution 134,224,98] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
250. 2 | ~18 [avatar contradiction clause 246]
262. ~r2_hidden(X3,X4) | v1_xtuple_0(X3) | ~v1_relat_1(X4) [superposition 111,102] {a:1,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
266. r2_hidden(k2_xtuple_0(sK2),sK1) <- (15) [resolution 209,108] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
267. r2_hidden(k1_xtuple_0(sK2),sK0) <- (15) [resolution 209,107] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
272. 19 <=> r2_hidden(k2_xtuple_0(sK2),sK1) [avatar definition]
274. r2_hidden(k2_xtuple_0(sK2),sK1) <- (19) [avatar component clause 272] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
275. 19 | ~15 [avatar split clause 266,207,272]
277. 20 <=> r2_hidden(k1_xtuple_0(sK2),sK0) [avatar definition]
279. r2_hidden(k1_xtuple_0(sK2),sK0) <- (20) [avatar component clause 277] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
280. 20 | ~15 [avatar split clause 267,207,277]
286. ~m1_subset_1(k2_xtuple_0(sK2),sK1) | ~m1_subset_1(k1_xtuple_0(sK2),sK0) | ~v1_xtuple_0(sK2) [equality resolution 237] {a:2,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
288. 22 <=> v1_xtuple_0(sK2) [avatar definition]
290. ~v1_xtuple_0(sK2) <- (~22) [avatar component clause 288] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
292. 23 <=> m1_subset_1(k1_xtuple_0(sK2),sK0) [avatar definition]
296. 24 <=> m1_subset_1(k2_xtuple_0(sK2),sK1) [avatar definition]
299. ~22 | ~23 | ~24 [avatar split clause 286,296,292,288]
306. m1_subset_1(k2_xtuple_0(sK2),sK1) <- (19) [resolution 274,104] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
313. 24 | ~19 [avatar split clause 306,272,296]
317. m1_subset_1(k1_xtuple_0(sK2),sK0) <- (20) [resolution 279,104] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
324. 23 | ~20 [avatar split clause 317,277,292]
334. v1_xtuple_0(sK2) | ~v1_relat_1(k2_zfmisc_1(sK0,sK1)) <- (15) [resolution 262,209] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
337. ~v1_relat_1(k2_zfmisc_1(sK0,sK1)) <- (15, ~22) [subsumption resolution 334,290] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
338. $false <- (15, ~22) [subsumption resolution 337,112] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
339. ~15 | 22 [avatar contradiction clause 338]
340. $false [avatar sat refutation 130,135,140,214,225,232,250,275,280,299,313,324,339] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:202,thDist:-202}
% SZS output end Proof for mcart_1__l23_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 41
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 38
% Generated clauses: 143
% Active clauses: 55
% Passive clauses: 79
% Final active clauses: 51
% Final passive clauses: 22
% Main loop iterations started: 56
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 14
% Fw demodulations to eq. taut.: 4
% 
% >>> Generating Inferences
% Binary resolution: 29
% Unit resulting resolution: 7
% Forward superposition: 12
% Backward superposition: 17
% Self superposition: 2
% Equality resolution: 2
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 10
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 35
% SAT solver unit clauses: 12
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 5117
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% superposition: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

