i: [1,128,0,2,1,1,0,0] t38_tsep_2
a: [3,204,0,2,1,128]
s: 204
i: [1,126,0,2,1,1,0,0] t38_tsep_2
a: [3,214,0,2,1,126]
s: 214
i: [1,124,0,2,1,1,0,0] t38_tsep_2
a: [3,224,0,2,1,124]
s: 224
i: [1,122,0,2,1,1,0,0] t38_tsep_2
a: [3,234,0,2,1,122]
s: 234
i: [1,121,0,2,1,1,0,0] t38_tsep_2
a: [3,239,0,2,1,121]
s: 239
i: [1,120,0,2,1,1,0,0] t38_tsep_2
a: [3,244,0,2,1,120]
s: 244
i: [1,119,0,2,1,1,0,0] t38_tsep_2
a: [3,249,0,2,1,119]
s: 249
i: [1,129,0,3,1,1,0,0] t38_tsep_2
a: [3,199,0,3,1,129]
s: 199
i: [1,127,0,3,1,1,0,0] t38_tsep_2
a: [3,209,0,3,1,127]
s: 209
i: [1,125,0,3,1,1,0,0] t38_tsep_2
a: [3,219,0,3,1,125]
s: 219
i: [1,123,0,3,1,1,0,0] t38_tsep_2
a: [3,229,0,3,1,123]
s: 229
i: [1,133,0,4,1,1,0,0] t38_tsep_2
a: [3,179,0,4,1,133]
s: 179
i: [1,132,0,4,1,1,0,0] t38_tsep_2
a: [3,184,0,4,1,132]
s: 184
i: [1,131,0,4,1,1,0,0] t38_tsep_2
a: [3,189,0,4,1,131]
s: 189
i: [1,130,0,4,1,1,0,0] t38_tsep_2
a: [3,194,0,4,1,130]
s: 194
i: [1,134,0,2,1,0,0,255] spc1_boole
a: [3,254,0,2,1,134]
s: 254
i: [1,135,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,259,0,2,1,135]
s: 259
i: [1,174,0,2,1,0,0,255] existence_l1_struct_0
a: [3,269,0,2,1,174]
s: 269
i: [1,175,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,274,0,2,1,175]
s: 274
i: [1,136,0,3,1,0,0,255] d2_xboole_0
a: [3,264,0,3,1,136]
s: 264
r: 259
d: [2,276,0,2,1,2,23,259,264]
a: [3,280,0,2,1,276]
s: 280
i: [1,152,0,4,2,0,0,255] dt_l1_pre_topc
s: 152
i: [1,164,0,4,1,0,0,255] existence_m1_subset_1
s: 164
d: [2,282,1,2,1,1,40,152,239]
a: [3,287,1,2,1,282]
s: 287
i: [1,137,0,5,2,0,0,255] t6_boole
s: 137
i: [1,163,0,5,2,0,0,255] t7_boole
s: 163
d: [2,283,1,2,1,1,40,152,274]
a: [3,292,1,2,1,283]
s: 292
i: [1,138,0,6,3,0,0,255] cc7_struct_0
s: 138
i: [1,139,0,6,3,0,0,255] cc5_struct_0
s: 139
i: [1,141,0,6,3,0,0,255] cc17_struct_0
s: 141
i: [1,142,0,6,3,0,0,255] cc2_struct_0
s: 142
d: [2,296,1,4,2,1,40,139,269]
a: [3,302,1,2,1,296]
s: 302
i: [1,143,0,6,3,0,0,255] cc6_struct_0
s: 143
d: [2,298,2,4,2,1,40,139,292]
a: [3,311,2,2,1,298]
s: 311
i: [1,144,0,6,3,0,0,255] cc1_struct_0
s: 144
i: [1,147,0,6,3,0,0,255] cc4_struct_0
s: 147
i: [1,154,0,6,2,0,0,1] existence_m1_pre_topc
s: 154
i: [1,160,0,6,2,0,0,255] t1_subset
s: 160
i: [1,165,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 165
i: [1,145,0,7,3,0,0,255] cc8_struct_0
s: 145
i: [1,148,0,7,3,0,0,255] cc11_struct_0
s: 148
i: [1,149,0,7,3,0,0,255] cc11_struct_0
s: 149
i: [1,150,0,7,3,0,0,255] cc9_struct_0
s: 150
i: [1,153,0,7,3,0,0,1] dt_m1_pre_topc
s: 153
i: [1,162,0,7,3,0,0,255] t8_boole
s: 162
i: [1,161,0,8,3,0,0,255] t2_subset
s: 161
d: [2,334,1,4,2,1,40,153,229]
d: [2,340,1,2,1,2,21,334,239]
a: [3,344,1,2,1,340]
s: 344
i: [1,151,0,9,4,0,0,255] cc10_struct_0
s: 151
d: [2,335,1,4,2,1,40,153,219]
d: [2,346,1,2,1,2,21,335,239]
a: [3,350,1,2,1,346]
s: 350
i: [1,159,0,9,4,0,0,1] cc1_pre_topc
s: 159
i: [1,166,0,10,4,0,0,255] symmetry_r1_tsep_1
s: 166
d: [2,336,1,4,2,1,40,153,209]
d: [2,352,1,2,1,2,21,336,239]
a: [3,356,1,2,1,352]
s: 356
i: [1,172,0,16,5,0,0,1] symmetry_r3_tsep_1
s: 172
d: [2,337,1,4,2,1,40,153,199]
d: [2,358,1,2,1,2,21,337,239]
a: [3,362,1,2,1,358]
s: 362
i: [1,173,0,16,5,0,0,1] symmetry_r4_tsep_1
s: 173
i: [1,167,0,20,8,0,0,1] irreflexivity_r4_tsep_2
s: 167
d: [2,381,1,6,3,1,40,159,229]
d: [2,387,1,4,2,2,21,381,239]
d: [2,388,1,2,1,3,21,387,244]
a: [3,392,1,2,1,388]
s: 392
i: [1,171,0,20,7,0,0,255] symmetry_r3_tsep_2
s: 171
d: [2,382,1,6,3,1,40,159,219]
d: [2,394,1,4,2,2,21,382,239]
d: [2,395,1,2,1,3,21,394,244]
a: [3,399,1,2,1,395]
s: 399
i: [1,157,0,23,9,0,0,255] t78_tsep_1
s: 157
i: [1,158,0,24,9,0,0,255] t78_tsep_1
s: 158
d: [2,383,1,6,3,1,40,159,209]
d: [2,401,1,4,2,2,21,383,239]
d: [2,402,1,2,1,3,21,401,244]
a: [3,406,1,2,1,402]
s: 406
i: [1,168,0,24,9,0,0,1] symmetry_r4_tsep_2
s: 168
d: [2,384,1,6,3,1,40,159,199]
d: [2,408,1,4,2,2,21,384,239]
d: [2,409,1,2,1,3,21,408,244]
a: [3,413,1,2,1,409]
s: 413
i: [1,169,0,24,9,0,0,255] redefinition_r4_tsep_2
s: 169
i: [1,170,0,24,9,0,0,255] redefinition_r4_tsep_2
s: 170
d: [2,365,2,2,1,1,40,344,152]
a: [3,369,2,2,1,365]
s: 369
i: [1,156,0,27,10,0,0,255] t78_tsep_1
s: 156
d: [2,375,2,2,1,1,40,350,152]
a: [3,379,2,2,1,375]
s: 379
i: [1,155,0,34,11,0,0,255] t37_tsep_2
s: 155
d: [2,436,1,19,8,1,40,157,184]
d: [2,437,1,16,7,2,21,436,219]
d: [2,438,1,14,6,3,21,437,224]
d: [2,439,1,11,5,4,21,438,229]
d: [2,440,1,9,4,5,21,439,234]
d: [2,441,1,7,3,6,21,440,239]
d: [2,442,1,5,2,7,21,441,244]
d: [2,443,1,3,1,8,21,442,249]
a: [3,447,1,3,1,443]
s: 447
d: [2,415,2,2,1,1,40,356,152]
a: [3,419,2,2,1,415]
s: 419
d: [2,421,1,12,4,1,40,172,184]
d: [2,422,1,9,3,2,21,421,219]
d: [2,423,1,6,2,3,21,422,229]
d: [2,424,1,4,1,4,21,423,239]
a: [3,428,1,4,1,424]
s: 428
d: [2,430,2,2,1,1,40,362,152]
a: [3,434,2,2,1,430]
s: 434
d: [2,449,1,20,8,1,40,158,184]
d: [2,450,1,17,7,2,21,449,219]
d: [2,451,1,15,6,3,21,450,224]
d: [2,452,1,12,5,4,21,451,229]
d: [2,453,1,10,4,5,21,452,234]
d: [2,454,1,8,3,6,21,453,239]
d: [2,455,1,6,2,7,21,454,244]
d: [2,456,1,4,1,8,21,455,249]
a: [3,460,1,4,1,456]
s: 460
d: [2,462,1,20,8,1,40,168,189]
d: [2,464,1,17,7,2,21,462,199]
d: [2,465,1,15,6,3,21,464,204]
d: [2,466,1,12,5,4,21,465,219]
d: [2,467,1,10,4,5,21,466,224]
d: [2,468,1,8,3,6,21,467,239]
d: [2,469,1,6,2,7,21,468,244]
d: [2,470,1,4,1,8,21,469,249]
a: [3,474,1,4,1,470]
s: 474
d: [2,518,2,7,3,1,40,447,166]
d: [2,519,2,5,2,2,21,518,379]
d: [2,520,2,3,1,3,21,519,369]
a: [3,524,2,3,1,520]
s: 524
d: [2,463,1,20,8,1,40,168,194]
d: [2,476,1,17,7,2,21,463,209]
d: [2,477,1,15,6,3,21,476,214]
d: [2,478,1,12,5,4,21,477,229]
d: [2,479,1,10,4,5,21,478,234]
d: [2,480,1,8,3,6,21,479,239]
d: [2,481,1,6,2,7,21,480,244]
d: [2,482,1,4,1,8,21,481,249]
a: [3,486,1,4,1,482]
s: 486
d: [2,488,1,20,8,1,40,169,189]
d: [2,490,1,17,7,2,21,488,199]
d: [2,491,1,15,6,3,21,490,204]
d: [2,492,1,12,5,4,21,491,219]
d: [2,493,1,10,4,5,21,492,224]
d: [2,494,1,8,3,6,21,493,239]
d: [2,495,1,6,2,7,21,494,244]
d: [2,496,1,4,1,8,21,495,249]
a: [3,500,1,4,1,496]
s: 500
d: [2,489,1,20,8,1,40,169,194]
d: [2,502,1,17,7,2,21,489,209]
d: [2,503,1,15,6,3,21,502,214]
d: [2,504,1,12,5,4,21,503,229]
d: [2,505,1,10,4,5,21,504,234]
d: [2,506,1,8,3,6,21,505,239]
d: [2,507,1,6,2,7,21,506,244]
d: [2,508,1,4,1,8,21,507,249]
a: [3,512,1,4,1,508]
s: 512
d: [2,602,2,0,0,11,52,249,244,239,229,219,209,199,500,179,460,512,155]
e: 602
f: 602,508,496,456,119,120,121,122,123,124,125,126,127,128,129,130,131,132,133
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tsep_2__t38_tsep_2
% SZS output start Proof for tsep_2__t38_tsep_2
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : ((m1_pre_topc(X1,X0) & ~v2_struct_0(X1)) => ! [X2] : ((m1_pre_topc(X2,X0) & ~v2_struct_0(X2)) => ! [X3] : ((m1_pre_topc(X3,X0) & ~v2_struct_0(X3)) => ! [X4] : ((m1_pre_topc(X4,X0) & ~v2_struct_0(X4)) => ((r3_tsep_1(X0,X1,X2) & r4_tsep_2(X0,X2,X4) & r4_tsep_2(X0,X1,X3)) => r4_tsep_1(X0,X3,X4))))))) [input t38_tsep_2]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : ((m1_pre_topc(X1,X0) & ~v2_struct_0(X1)) => ! [X2] : ((m1_pre_topc(X2,X0) & ~v2_struct_0(X2)) => ! [X3] : ((m1_pre_topc(X3,X0) & ~v2_struct_0(X3)) => ! [X4] : ((m1_pre_topc(X4,X0) & ~v2_struct_0(X4)) => ((r3_tsep_1(X0,X1,X2) & r4_tsep_2(X0,X2,X4) & r4_tsep_2(X0,X1,X3)) => r4_tsep_1(X0,X3,X4))))))) [negated conjecture 1]
29. ! [X0,X1,X2] : ((m1_pre_topc(X2,X0) & ~v2_struct_0(X2) & m1_pre_topc(X1,X0) & ~v2_struct_0(X1) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (r4_tsep_2(X0,X1,X2) <=> r3_tsep_2(X0,X1,X2))) [input redefinition_r4_tsep_2]
38. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_pre_topc(X1,X0) => ! [X2] : (m1_pre_topc(X2,X0) => ! [X3] : (m1_pre_topc(X3,X0) => ! [X4] : (m1_pre_topc(X4,X0) => ((r4_tsep_1(X0,X1,X2) & r3_tsep_2(X0,X2,X4) & r3_tsep_2(X0,X1,X3)) => r4_tsep_1(X0,X3,X4))))))) [input t37_tsep_2]
40. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : ((m1_pre_topc(X1,X0) & ~v2_struct_0(X1)) => ! [X2] : ((m1_pre_topc(X2,X0) & ~v2_struct_0(X2)) => ((r4_tsep_1(X0,X1,X2) & r1_tsep_1(X1,X2)) <=> r3_tsep_1(X0,X1,X2))))) [input t78_tsep_1]
48. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : ((~r4_tsep_1(X0,X3,X4) & (r3_tsep_1(X0,X1,X2) & r4_tsep_2(X0,X2,X4) & r4_tsep_2(X0,X1,X3))) & (m1_pre_topc(X4,X0) & ~v2_struct_0(X4))) & (m1_pre_topc(X3,X0) & ~v2_struct_0(X3))) & (m1_pre_topc(X2,X0) & ~v2_struct_0(X2))) & (m1_pre_topc(X1,X0) & ~v2_struct_0(X1))) & (l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
49. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r4_tsep_1(X0,X3,X4) & r3_tsep_1(X0,X1,X2) & r4_tsep_2(X0,X2,X4) & r4_tsep_2(X0,X1,X3) & m1_pre_topc(X4,X0) & ~v2_struct_0(X4)) & m1_pre_topc(X3,X0) & ~v2_struct_0(X3)) & m1_pre_topc(X2,X0) & ~v2_struct_0(X2)) & m1_pre_topc(X1,X0) & ~v2_struct_0(X1)) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [flattening 48]
76. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r4_tsep_1(X0,X3,X4) | (~r4_tsep_1(X0,X1,X2) | ~r3_tsep_2(X0,X2,X4) | ~r3_tsep_2(X0,X1,X3))) | ~m1_pre_topc(X4,X0)) | ~m1_pre_topc(X3,X0)) | ~m1_pre_topc(X2,X0)) | ~m1_pre_topc(X1,X0)) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0))) [ennf transformation 38]
77. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (r4_tsep_1(X0,X3,X4) | ~r4_tsep_1(X0,X1,X2) | ~r3_tsep_2(X0,X2,X4) | ~r3_tsep_2(X0,X1,X3) | ~m1_pre_topc(X4,X0)) | ~m1_pre_topc(X3,X0)) | ~m1_pre_topc(X2,X0)) | ~m1_pre_topc(X1,X0)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [flattening 76]
78. ! [X0] : (! [X1] : (! [X2] : (((r4_tsep_1(X0,X1,X2) & r1_tsep_1(X1,X2)) <=> r3_tsep_1(X0,X1,X2)) | (~m1_pre_topc(X2,X0) | v2_struct_0(X2))) | (~m1_pre_topc(X1,X0) | v2_struct_0(X1))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0))) [ennf transformation 40]
79. ! [X0] : (! [X1] : (! [X2] : (((r4_tsep_1(X0,X1,X2) & r1_tsep_1(X1,X2)) <=> r3_tsep_1(X0,X1,X2)) | ~m1_pre_topc(X2,X0) | v2_struct_0(X2)) | ~m1_pre_topc(X1,X0) | v2_struct_0(X1)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [flattening 78]
94. ! [X0,X1,X2] : ((r4_tsep_2(X0,X1,X2) <=> r3_tsep_2(X0,X1,X2)) | (~m1_pre_topc(X2,X0) | v2_struct_0(X2) | ~m1_pre_topc(X1,X0) | v2_struct_0(X1) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0))) [ennf transformation 29]
95. ! [X0,X1,X2] : ((r4_tsep_2(X0,X1,X2) <=> r3_tsep_2(X0,X1,X2)) | ~m1_pre_topc(X2,X0) | v2_struct_0(X2) | ~m1_pre_topc(X1,X0) | v2_struct_0(X1) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [flattening 94]
102. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r4_tsep_1(X0,X3,X4) & r3_tsep_1(X0,X1,X2) & r4_tsep_2(X0,X2,X4) & r4_tsep_2(X0,X1,X3) & m1_pre_topc(X4,X0) & ~v2_struct_0(X4)) & m1_pre_topc(X3,X0) & ~v2_struct_0(X3)) & m1_pre_topc(X2,X0) & ~v2_struct_0(X2)) & m1_pre_topc(X1,X0) & ~v2_struct_0(X1)) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r4_tsep_1(sK0,X3,X4) & r3_tsep_1(sK0,X1,X2) & r4_tsep_2(sK0,X2,X4) & r4_tsep_2(sK0,X1,X3) & m1_pre_topc(X4,sK0) & ~v2_struct_0(X4)) & m1_pre_topc(X3,sK0) & ~v2_struct_0(X3)) & m1_pre_topc(X2,sK0) & ~v2_struct_0(X2)) & m1_pre_topc(X1,sK0) & ~v2_struct_0(X1)) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0)) [choice axiom]
103. ? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r4_tsep_1(sK0,X3,X4) & r3_tsep_1(sK0,X1,X2) & r4_tsep_2(sK0,X2,X4) & r4_tsep_2(sK0,X1,X3) & m1_pre_topc(X4,sK0) & ~v2_struct_0(X4)) & m1_pre_topc(X3,sK0) & ~v2_struct_0(X3)) & m1_pre_topc(X2,sK0) & ~v2_struct_0(X2)) & m1_pre_topc(X1,sK0) & ~v2_struct_0(X1)) => (? [X2] : (? [X3] : (? [X4] : (~r4_tsep_1(sK0,X3,X4) & r3_tsep_1(sK0,sK1,X2) & r4_tsep_2(sK0,X2,X4) & r4_tsep_2(sK0,sK1,X3) & m1_pre_topc(X4,sK0) & ~v2_struct_0(X4)) & m1_pre_topc(X3,sK0) & ~v2_struct_0(X3)) & m1_pre_topc(X2,sK0) & ~v2_struct_0(X2)) & m1_pre_topc(sK1,sK0) & ~v2_struct_0(sK1)) [choice axiom]
104. ? [X2] : (? [X3] : (? [X4] : (~r4_tsep_1(sK0,X3,X4) & r3_tsep_1(sK0,sK1,X2) & r4_tsep_2(sK0,X2,X4) & r4_tsep_2(sK0,sK1,X3) & m1_pre_topc(X4,sK0) & ~v2_struct_0(X4)) & m1_pre_topc(X3,sK0) & ~v2_struct_0(X3)) & m1_pre_topc(X2,sK0) & ~v2_struct_0(X2)) => (? [X3] : (? [X4] : (~r4_tsep_1(sK0,X3,X4) & r3_tsep_1(sK0,sK1,sK2) & r4_tsep_2(sK0,sK2,X4) & r4_tsep_2(sK0,sK1,X3) & m1_pre_topc(X4,sK0) & ~v2_struct_0(X4)) & m1_pre_topc(X3,sK0) & ~v2_struct_0(X3)) & m1_pre_topc(sK2,sK0) & ~v2_struct_0(sK2)) [choice axiom]
105. ? [X3] : (? [X4] : (~r4_tsep_1(sK0,X3,X4) & r3_tsep_1(sK0,sK1,sK2) & r4_tsep_2(sK0,sK2,X4) & r4_tsep_2(sK0,sK1,X3) & m1_pre_topc(X4,sK0) & ~v2_struct_0(X4)) & m1_pre_topc(X3,sK0) & ~v2_struct_0(X3)) => (? [X4] : (~r4_tsep_1(sK0,sK3,X4) & r3_tsep_1(sK0,sK1,sK2) & r4_tsep_2(sK0,sK2,X4) & r4_tsep_2(sK0,sK1,sK3) & m1_pre_topc(X4,sK0) & ~v2_struct_0(X4)) & m1_pre_topc(sK3,sK0) & ~v2_struct_0(sK3)) [choice axiom]
106. ? [X4] : (~r4_tsep_1(sK0,sK3,X4) & r3_tsep_1(sK0,sK1,sK2) & r4_tsep_2(sK0,sK2,X4) & r4_tsep_2(sK0,sK1,sK3) & m1_pre_topc(X4,sK0) & ~v2_struct_0(X4)) => (~r4_tsep_1(sK0,sK3,sK4) & r3_tsep_1(sK0,sK1,sK2) & r4_tsep_2(sK0,sK2,sK4) & r4_tsep_2(sK0,sK1,sK3) & m1_pre_topc(sK4,sK0) & ~v2_struct_0(sK4)) [choice axiom]
107. ((((~r4_tsep_1(sK0,sK3,sK4) & r3_tsep_1(sK0,sK1,sK2) & r4_tsep_2(sK0,sK2,sK4) & r4_tsep_2(sK0,sK1,sK3) & m1_pre_topc(sK4,sK0) & ~v2_struct_0(sK4)) & m1_pre_topc(sK3,sK0) & ~v2_struct_0(sK3)) & m1_pre_topc(sK2,sK0) & ~v2_struct_0(sK2)) & m1_pre_topc(sK1,sK0) & ~v2_struct_0(sK1)) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0) [skolemisation 49,106,105,104,103,102]
110. ! [X0] : (! [X1] : (! [X2] : ((((r4_tsep_1(X0,X1,X2) & r1_tsep_1(X1,X2)) | ~r3_tsep_1(X0,X1,X2)) & (r3_tsep_1(X0,X1,X2) | (~r4_tsep_1(X0,X1,X2) | ~r1_tsep_1(X1,X2)))) | ~m1_pre_topc(X2,X0) | v2_struct_0(X2)) | ~m1_pre_topc(X1,X0) | v2_struct_0(X1)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [nnf transformation 79]
111. ! [X0] : (! [X1] : (! [X2] : ((((r4_tsep_1(X0,X1,X2) & r1_tsep_1(X1,X2)) | ~r3_tsep_1(X0,X1,X2)) & (r3_tsep_1(X0,X1,X2) | ~r4_tsep_1(X0,X1,X2) | ~r1_tsep_1(X1,X2))) | ~m1_pre_topc(X2,X0) | v2_struct_0(X2)) | ~m1_pre_topc(X1,X0) | v2_struct_0(X1)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [flattening 110]
114. ! [X0,X1,X2] : (((r4_tsep_2(X0,X1,X2) | ~r3_tsep_2(X0,X1,X2)) & (r3_tsep_2(X0,X1,X2) | ~r4_tsep_2(X0,X1,X2))) | ~m1_pre_topc(X2,X0) | v2_struct_0(X2) | ~m1_pre_topc(X1,X0) | v2_struct_0(X1) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [nnf transformation 95]
119. ~v2_struct_0(sK0) [cnf transformation 107] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
120. v2_pre_topc(sK0) [cnf transformation 107] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
121. l1_pre_topc(sK0) [cnf transformation 107] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
122. ~v2_struct_0(sK1) [cnf transformation 107] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
123. m1_pre_topc(sK1,sK0) [cnf transformation 107] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
124. ~v2_struct_0(sK2) [cnf transformation 107] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
125. m1_pre_topc(sK2,sK0) [cnf transformation 107] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
126. ~v2_struct_0(sK3) [cnf transformation 107] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
127. m1_pre_topc(sK3,sK0) [cnf transformation 107] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
128. ~v2_struct_0(sK4) [cnf transformation 107] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
129. m1_pre_topc(sK4,sK0) [cnf transformation 107] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
130. r4_tsep_2(sK0,sK1,sK3) [cnf transformation 107] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
131. r4_tsep_2(sK0,sK2,sK4) [cnf transformation 107] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
132. r3_tsep_1(sK0,sK1,sK2) [cnf transformation 107] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
133. ~r4_tsep_1(sK0,sK3,sK4) [cnf transformation 107] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
155. ~r3_tsep_2(X0,X2,X4) | ~r4_tsep_1(X0,X1,X2) | r4_tsep_1(X0,X3,X4) | ~r3_tsep_2(X0,X1,X3) | ~m1_pre_topc(X4,X0) | ~m1_pre_topc(X3,X0) | ~m1_pre_topc(X2,X0) | ~m1_pre_topc(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 77] {a:0,w:34,wCS:13600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
158. ~r3_tsep_1(X0,X1,X2) | r4_tsep_1(X0,X1,X2) | ~m1_pre_topc(X2,X0) | v2_struct_0(X2) | ~m1_pre_topc(X1,X0) | v2_struct_0(X1) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 111] {a:0,w:24,wCS:9600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
169. ~r4_tsep_2(X0,X1,X2) | r3_tsep_2(X0,X1,X2) | ~m1_pre_topc(X2,X0) | v2_struct_0(X2) | ~m1_pre_topc(X1,X0) | v2_struct_0(X1) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 114] {a:0,w:24,wCS:9600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
177. 1 <=> r4_tsep_1(sK0,sK3,sK4) [avatar definition]
179. ~r4_tsep_1(sK0,sK3,sK4) <- (~1) [avatar component clause 177] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
180. ~1 [avatar split clause 133,177]
182. 2 <=> r3_tsep_1(sK0,sK1,sK2) [avatar definition]
184. r3_tsep_1(sK0,sK1,sK2) <- (2) [avatar component clause 182] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
185. 2 [avatar split clause 132,182]
187. 3 <=> r4_tsep_2(sK0,sK2,sK4) [avatar definition]
189. r4_tsep_2(sK0,sK2,sK4) <- (3) [avatar component clause 187] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
190. 3 [avatar split clause 131,187]
192. 4 <=> r4_tsep_2(sK0,sK1,sK3) [avatar definition]
194. r4_tsep_2(sK0,sK1,sK3) <- (4) [avatar component clause 192] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
195. 4 [avatar split clause 130,192]
197. 5 <=> m1_pre_topc(sK4,sK0) [avatar definition]
199. m1_pre_topc(sK4,sK0) <- (5) [avatar component clause 197] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
200. 5 [avatar split clause 129,197]
202. 6 <=> v2_struct_0(sK4) [avatar definition]
204. ~v2_struct_0(sK4) <- (~6) [avatar component clause 202] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
205. ~6 [avatar split clause 128,202]
207. 7 <=> m1_pre_topc(sK3,sK0) [avatar definition]
209. m1_pre_topc(sK3,sK0) <- (7) [avatar component clause 207] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
210. 7 [avatar split clause 127,207]
212. 8 <=> v2_struct_0(sK3) [avatar definition]
214. ~v2_struct_0(sK3) <- (~8) [avatar component clause 212] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
215. ~8 [avatar split clause 126,212]
217. 9 <=> m1_pre_topc(sK2,sK0) [avatar definition]
219. m1_pre_topc(sK2,sK0) <- (9) [avatar component clause 217] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
220. 9 [avatar split clause 125,217]
222. 10 <=> v2_struct_0(sK2) [avatar definition]
224. ~v2_struct_0(sK2) <- (~10) [avatar component clause 222] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
225. ~10 [avatar split clause 124,222]
227. 11 <=> m1_pre_topc(sK1,sK0) [avatar definition]
229. m1_pre_topc(sK1,sK0) <- (11) [avatar component clause 227] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
230. 11 [avatar split clause 123,227]
232. 12 <=> v2_struct_0(sK1) [avatar definition]
234. ~v2_struct_0(sK1) <- (~12) [avatar component clause 232] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
235. ~12 [avatar split clause 122,232]
237. 13 <=> l1_pre_topc(sK0) [avatar definition]
239. l1_pre_topc(sK0) <- (13) [avatar component clause 237] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
240. 13 [avatar split clause 121,237]
242. 14 <=> v2_pre_topc(sK0) [avatar definition]
244. v2_pre_topc(sK0) <- (14) [avatar component clause 242] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
245. 14 [avatar split clause 120,242]
247. 15 <=> v2_struct_0(sK0) [avatar definition]
249. ~v2_struct_0(sK0) <- (~15) [avatar component clause 247] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
250. ~15 [avatar split clause 119,247]
449. r4_tsep_1(sK0,sK1,sK2) | ~m1_pre_topc(sK2,sK0) | v2_struct_0(sK2) | ~m1_pre_topc(sK1,sK0) | v2_struct_0(sK1) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2) [resolution 158,184] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
450. r4_tsep_1(sK0,sK1,sK2) | v2_struct_0(sK2) | ~m1_pre_topc(sK1,sK0) | v2_struct_0(sK1) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2, 9) [subsumption resolution 449,219] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
451. r4_tsep_1(sK0,sK1,sK2) | ~m1_pre_topc(sK1,sK0) | v2_struct_0(sK1) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2, 9, ~10) [subsumption resolution 450,224] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
452. r4_tsep_1(sK0,sK1,sK2) | v2_struct_0(sK1) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2, 9, ~10, 11) [subsumption resolution 451,229] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
453. r4_tsep_1(sK0,sK1,sK2) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2, 9, ~10, 11, ~12) [subsumption resolution 452,234] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
454. r4_tsep_1(sK0,sK1,sK2) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (2, 9, ~10, 11, ~12, 13) [subsumption resolution 453,239] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
455. r4_tsep_1(sK0,sK1,sK2) | v2_struct_0(sK0) <- (2, 9, ~10, 11, ~12, 13, 14) [subsumption resolution 454,244] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
456. r4_tsep_1(sK0,sK1,sK2) <- (2, 9, ~10, 11, ~12, 13, 14, ~15) [subsumption resolution 455,249] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
458. 44 <=> r4_tsep_1(sK0,sK1,sK2) [avatar definition]
460. r4_tsep_1(sK0,sK1,sK2) <- (44) [avatar component clause 458] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
461. 44 | ~2 | ~9 | 10 | ~11 | 12 | ~13 | ~14 | 15 [avatar split clause 456,247,242,237,232,227,222,217,182,458]
488. r3_tsep_2(sK0,sK2,sK4) | ~m1_pre_topc(sK4,sK0) | v2_struct_0(sK4) | ~m1_pre_topc(sK2,sK0) | v2_struct_0(sK2) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (3) [resolution 169,189] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
489. r3_tsep_2(sK0,sK1,sK3) | ~m1_pre_topc(sK3,sK0) | v2_struct_0(sK3) | ~m1_pre_topc(sK1,sK0) | v2_struct_0(sK1) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4) [resolution 169,194] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
490. r3_tsep_2(sK0,sK2,sK4) | v2_struct_0(sK4) | ~m1_pre_topc(sK2,sK0) | v2_struct_0(sK2) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (3, 5) [subsumption resolution 488,199] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
491. r3_tsep_2(sK0,sK2,sK4) | ~m1_pre_topc(sK2,sK0) | v2_struct_0(sK2) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (3, 5, ~6) [subsumption resolution 490,204] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
492. r3_tsep_2(sK0,sK2,sK4) | v2_struct_0(sK2) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (3, 5, ~6, 9) [subsumption resolution 491,219] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
493. r3_tsep_2(sK0,sK2,sK4) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (3, 5, ~6, 9, ~10) [subsumption resolution 492,224] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
494. r3_tsep_2(sK0,sK2,sK4) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (3, 5, ~6, 9, ~10, 13) [subsumption resolution 493,239] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
495. r3_tsep_2(sK0,sK2,sK4) | v2_struct_0(sK0) <- (3, 5, ~6, 9, ~10, 13, 14) [subsumption resolution 494,244] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
496. r3_tsep_2(sK0,sK2,sK4) <- (3, 5, ~6, 9, ~10, 13, 14, ~15) [subsumption resolution 495,249] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
498. 47 <=> r3_tsep_2(sK0,sK2,sK4) [avatar definition]
500. r3_tsep_2(sK0,sK2,sK4) <- (47) [avatar component clause 498] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
501. 47 | ~3 | ~5 | 6 | ~9 | 10 | ~13 | ~14 | 15 [avatar split clause 496,247,242,237,222,217,202,197,187,498]
502. r3_tsep_2(sK0,sK1,sK3) | v2_struct_0(sK3) | ~m1_pre_topc(sK1,sK0) | v2_struct_0(sK1) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 7) [subsumption resolution 489,209] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
503. r3_tsep_2(sK0,sK1,sK3) | ~m1_pre_topc(sK1,sK0) | v2_struct_0(sK1) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 7, ~8) [subsumption resolution 502,214] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
504. r3_tsep_2(sK0,sK1,sK3) | v2_struct_0(sK1) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 7, ~8, 11) [subsumption resolution 503,229] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
505. r3_tsep_2(sK0,sK1,sK3) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 7, ~8, 11, ~12) [subsumption resolution 504,234] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
506. r3_tsep_2(sK0,sK1,sK3) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 7, ~8, 11, ~12, 13) [subsumption resolution 505,239] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
507. r3_tsep_2(sK0,sK1,sK3) | v2_struct_0(sK0) <- (4, 7, ~8, 11, ~12, 13, 14) [subsumption resolution 506,244] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
508. r3_tsep_2(sK0,sK1,sK3) <- (4, 7, ~8, 11, ~12, 13, 14, ~15) [subsumption resolution 507,249] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
510. 48 <=> r3_tsep_2(sK0,sK1,sK3) [avatar definition]
512. r3_tsep_2(sK0,sK1,sK3) <- (48) [avatar component clause 510] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
513. 48 | ~4 | ~7 | 8 | ~11 | 12 | ~13 | ~14 | 15 [avatar split clause 508,247,242,237,232,227,212,207,192,510]
602. $false <- (~1, 5, 7, 9, 11, 13, 14, ~15, 44, 47, 48) [unit resulting resolution 249,244,239,229,219,209,199,500,179,460,512,155] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:70,thDist:-70}
606. 1 | ~5 | ~7 | ~9 | ~11 | ~13 | ~14 | 15 | ~44 | ~47 | ~48 [avatar contradiction clause 602]
617. $false [avatar sat refutation 180,185,190,195,200,205,210,215,220,225,230,235,240,245,250,461,501,513,606] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:571,thDist:-571}
% SZS output end Proof for tsep_2__t38_tsep_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 41
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 57
% Generated clauses: 284
% Active clauses: 76
% Passive clauses: 90
% Final active clauses: 75
% Final passive clauses: 10
% Main loop iterations started: 80
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Fw subsumption resolutions: 110
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Forward subsumptions: 21
% 
% >>> Generating Inferences
% Binary resolution: 58
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 56
% SAT solver unit clauses: 20
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 9722
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% backward subsumption index maintenance: 0.001 s
% resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

