i: [1,46,0,2,1,1,0,0] t45_algspec1
a: [3,66,0,2,1,46]
s: 66
i: [1,45,0,2,1,1,0,0] t45_algspec1
a: [3,71,0,2,1,45]
s: 71
i: [1,47,0,3,1,1,0,0] t45_algspec1
a: [3,61,0,3,1,47]
s: 61
i: [1,57,0,2,1,0,0,1] existence_l1_msualg_1
a: [3,76,0,2,1,57]
s: 76
i: [1,48,0,7,3,0,0,2] t15_instalg1
s: 48
i: [1,55,0,8,3,0,0,1] existence_m2_algspec1
s: 55
i: [1,56,0,8,3,0,0,2] existence_m1_instalg1
s: 56
i: [1,49,0,9,4,0,0,1] dt_m2_algspec1
s: 49
d: [2,78,1,11,5,0,40,49,55]
d: [2,79,1,7,3,0,18,78]
s: 79
i: [1,50,0,9,4,0,0,1] dt_m2_algspec1
s: 50
i: [1,51,0,9,4,0,0,2] cc6_instalg1
s: 51
d: [2,80,1,11,5,0,40,50,55]
d: [2,81,1,7,3,0,18,80]
s: 81
i: [1,52,0,9,4,0,0,2] dt_m1_instalg1
s: 52
d: [2,83,1,11,5,0,40,51,56]
d: [2,84,1,7,3,0,18,83]
s: 84
i: [1,53,0,14,6,0,0,1] d5_algspec1
s: 53
d: [2,90,0,12,5,0,21,53,49]
d: [2,91,0,10,4,0,21,90,50]
s: 91
d: [2,87,1,11,5,0,40,52,56]
d: [2,88,1,7,3,0,18,87]
s: 88
i: [1,54,0,14,6,0,0,1] d5_algspec1
s: 54
d: [2,92,1,12,5,0,40,91,55]
d: [2,93,1,8,3,0,18,92]
s: 93
d: [2,94,0,12,5,0,21,54,51]
d: [2,95,0,10,4,0,21,94,52]
s: 95
d: [2,98,1,11,5,0,40,95,48]
d: [2,102,1,7,3,0,18,98]
s: 102
d: [2,103,2,0,0,3,52,71,66,61,102]
e: 103
d: [2,104,2,4,2,1,40,102,61]
d: [2,112,2,2,1,2,21,104,66]
d: [2,113,2,0,0,3,21,112,71]
e: 113
f: 113,45,46,47
% Refutation found. Thanks to Tanya!
% SZS status Theorem for algspec1__t45_algspec1
% SZS output start Proof for algspec1__t45_algspec1
1. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => m2_algspec1(X0,X0)) [input t45_algspec1]
2. ~! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => m2_algspec1(X0,X0)) [negated conjecture 1]
3. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : (m1_instalg1(X1,X0) => v1_instalg1(X1))) [input cc6_instalg1]
4. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : ((l1_msualg_1(X1) & v1_instalg1(X1)) => (m2_algspec1(X1,X0) <=> m1_instalg1(X0,X1)))) [input d5_algspec1]
8. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => ! [X1] : (m1_instalg1(X1,X0) => l1_msualg_1(X1))) [input dt_m1_instalg1]
15. ! [X0] : ((l1_msualg_1(X0) & v1_instalg1(X0)) => m1_instalg1(X0,X0)) [input t15_instalg1]
20. ? [X0] : (~m2_algspec1(X0,X0) & (l1_msualg_1(X0) & v1_instalg1(X0))) [ennf transformation 2]
21. ? [X0] : (~m2_algspec1(X0,X0) & l1_msualg_1(X0) & v1_instalg1(X0)) [flattening 20]
22. ! [X0] : (m1_instalg1(X0,X0) | (~l1_msualg_1(X0) | ~v1_instalg1(X0))) [ennf transformation 15]
23. ! [X0] : (m1_instalg1(X0,X0) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [flattening 22]
26. ! [X0] : (! [X1] : (v1_instalg1(X1) | ~m1_instalg1(X1,X0)) | (~l1_msualg_1(X0) | ~v1_instalg1(X0))) [ennf transformation 3]
27. ! [X0] : (! [X1] : (v1_instalg1(X1) | ~m1_instalg1(X1,X0)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [flattening 26]
28. ! [X0] : (! [X1] : (l1_msualg_1(X1) | ~m1_instalg1(X1,X0)) | (~l1_msualg_1(X0) | ~v1_instalg1(X0))) [ennf transformation 8]
29. ! [X0] : (! [X1] : (l1_msualg_1(X1) | ~m1_instalg1(X1,X0)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [flattening 28]
30. ! [X0] : (! [X1] : ((m2_algspec1(X1,X0) <=> m1_instalg1(X0,X1)) | (~l1_msualg_1(X1) | ~v1_instalg1(X1))) | (~l1_msualg_1(X0) | ~v1_instalg1(X0))) [ennf transformation 4]
31. ! [X0] : (! [X1] : ((m2_algspec1(X1,X0) <=> m1_instalg1(X0,X1)) | ~l1_msualg_1(X1) | ~v1_instalg1(X1)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [flattening 30]
36. ? [X0] : (~m2_algspec1(X0,X0) & l1_msualg_1(X0) & v1_instalg1(X0)) => (~m2_algspec1(sK0,sK0) & l1_msualg_1(sK0) & v1_instalg1(sK0)) [choice axiom]
37. ~m2_algspec1(sK0,sK0) & l1_msualg_1(sK0) & v1_instalg1(sK0) [skolemisation 21,36]
38. ! [X0] : (! [X1] : (((m2_algspec1(X1,X0) | ~m1_instalg1(X0,X1)) & (m1_instalg1(X0,X1) | ~m2_algspec1(X1,X0))) | ~l1_msualg_1(X1) | ~v1_instalg1(X1)) | ~l1_msualg_1(X0) | ~v1_instalg1(X0)) [nnf transformation 31]
45. v1_instalg1(sK0) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. l1_msualg_1(sK0) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. ~m2_algspec1(sK0,sK0) [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. m1_instalg1(X0,X0) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 23] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
51. ~m1_instalg1(X1,X0) | v1_instalg1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 27] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
52. ~m1_instalg1(X1,X0) | l1_msualg_1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 29] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
54. m2_algspec1(X1,X0) | ~m1_instalg1(X0,X1) | ~l1_msualg_1(X1) | ~v1_instalg1(X1) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [cnf transformation 38] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:1,thDist:-1}
59. 1 <=> m2_algspec1(sK0,sK0) [avatar definition]
61. ~m2_algspec1(sK0,sK0) <- (~1) [avatar component clause 59] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. ~1 [avatar split clause 47,59]
64. 2 <=> l1_msualg_1(sK0) [avatar definition]
66. l1_msualg_1(sK0) <- (2) [avatar component clause 64] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
67. 2 [avatar split clause 46,64]
69. 3 <=> v1_instalg1(sK0) [avatar definition]
71. v1_instalg1(sK0) <- (3) [avatar component clause 69] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. 3 [avatar split clause 45,69]
94. m2_algspec1(X1,X0) | ~m1_instalg1(X0,X1) | ~l1_msualg_1(X1) | ~v1_instalg1(X1) | ~l1_msualg_1(X0) [subsumption resolution 54,51] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
95. ~m1_instalg1(X0,X1) | m2_algspec1(X1,X0) | ~l1_msualg_1(X1) | ~v1_instalg1(X1) [subsumption resolution 94,52] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
98. m2_algspec1(X0,X0) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [resolution 95,48] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
102. m2_algspec1(X0,X0) | ~l1_msualg_1(X0) | ~v1_instalg1(X0) [duplicate literal removal 98] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
104. ~l1_msualg_1(sK0) | ~v1_instalg1(sK0) <- (~1) [resolution 102,61] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
112. ~v1_instalg1(sK0) <- (~1, 2) [subsumption resolution 104,66] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
113. $false <- (~1, 2, 3) [subsumption resolution 112,71] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
114. 1 | ~2 | ~3 [avatar contradiction clause 113]
115. $false [avatar sat refutation 62,67,72,114] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:24,thDist:-24}
% SZS output end Proof for algspec1__t45_algspec1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 52
% Active clauses: 19
% Passive clauses: 22
% Final active clauses: 19
% Final passive clauses: 1
% Main loop iterations started: 21
% 
% >>> Simplifying Inferences
% Duplicate literals: 24
% Fw subsumption resolutions: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Forward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 16
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 8
% SAT solver unit clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.0000 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.0000 s

