i: [1,47,0,3,1,1,0,0] t67_zfmisc_1
a: [3,77,0,3,1,47]
s: 77
i: [1,48,0,5,1,1,0,0] t67_zfmisc_1
a: [3,72,0,5,1,48]
s: 72
i: [1,49,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,82,0,2,1,49]
s: 82
i: [1,50,0,2,1,0,0,255] fc1_xboole_0
a: [3,87,0,2,1,50]
s: 87
i: [1,65,0,2,1,0,0,3] rc2_xboole_0
a: [3,97,0,2,1,65]
s: 97
i: [1,66,0,2,1,0,0,3] rc1_xboole_0
a: [3,102,0,2,1,66]
s: 102
i: [1,63,0,3,1,0,0,1] reflexivity_r1_tarski
s: 63
i: [1,51,0,3,1,0,0,255] d2_xboole_0
a: [3,92,0,3,1,51]
s: 92
r: 82
i: [1,52,0,5,2,0,0,255] t6_boole
s: 52
i: [1,61,0,5,2,0,0,2] t7_boole
s: 61
i: [1,64,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 64
i: [1,60,0,7,3,0,0,2] t8_boole
s: 60
i: [1,67,0,7,2,0,0,1] d1_zfmisc_1
s: 67
d: [2,108,1,6,2,0,40,67,61]
s: 108
i: [1,68,0,7,2,0,0,1] d1_zfmisc_1
s: 68
i: [1,54,0,8,2,0,0,2] d3_tarski
s: 54
d: [2,111,1,5,2,0,40,54,61]
s: 111
i: [1,55,0,8,2,0,0,2] d3_tarski
s: 55
d: [2,113,2,3,1,1,40,111,72]
a: [3,117,2,3,1,113]
s: 117
i: [1,53,0,9,3,0,0,2] d3_tarski
s: 53
i: [1,62,0,9,3,0,0,1] t1_xboole_1
s: 62
d: [2,124,1,6,2,1,40,62,77]
s: 124
i: [1,58,0,14,3,0,0,1] d1_zfmisc_1
s: 58
d: [2,105,1,7,3,0,44,52,52]
s: 105
i: [1,59,0,14,3,0,0,1] d1_zfmisc_1
s: 59
d: [2,107,1,7,2,0,40,67,64]
s: 107
d: [2,110,1,8,2,0,40,54,64]
s: 110
d: [2,112,1,10,2,0,40,54,68]
s: 112
d: [2,138,2,10,2,1,40,112,124]
s: 138
d: [2,120,1,10,2,0,40,55,67]
s: 120
d: [2,141,3,10,2,1,40,120,138]
d: [2,143,3,5,1,1,18,141]
d: [2,145,3,0,0,2,21,143,72]
e: 145
f: 145,47,48
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t67_zfmisc_1
% SZS output start Proof for zfmisc_1__t67_zfmisc_1
1. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k1_zfmisc_1(X0),k1_zfmisc_1(X1))) [input t67_zfmisc_1]
2. ~! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k1_zfmisc_1(X0),k1_zfmisc_1(X1))) [negated conjecture 1]
4. ! [X0] : ! [X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [input d1_zfmisc_1]
6. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
14. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
18. ~! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k1_zfmisc_1(X0),k1_zfmisc_1(X1))) [flattening 2]
19. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 6]
20. ! [X0,X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [flattening 4]
23. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 14]
25. ? [X0,X1] : (~r1_tarski(k1_zfmisc_1(X0),k1_zfmisc_1(X1)) & r1_tarski(X0,X1)) [ennf transformation 18]
27. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 19]
30. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 23]
31. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 30]
33. ? [X0,X1] : (~r1_tarski(k1_zfmisc_1(X0),k1_zfmisc_1(X1)) & r1_tarski(X0,X1)) => (~r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) & r1_tarski(sK0,sK1)) [choice axiom]
34. ~r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) & r1_tarski(sK0,sK1) [skolemisation 25,33]
35. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 27]
36. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 35]
37. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) [choice axiom]
38. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 36,37]
39. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ~r1_tarski(X2,X0)) & (r1_tarski(X2,X0) | ~r2_hidden(X2,X1))) | k1_zfmisc_1(X0) != X1)) [nnf transformation 20]
40. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [rectify 39]
41. ! [X1,X0] : (? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1))) => ((~r1_tarski(sK3(X0,X1),X0) | ~r2_hidden(sK3(X0,X1),X1)) & (r1_tarski(sK3(X0,X1),X0) | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
42. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ((~r1_tarski(sK3(X0,X1),X0) | ~r2_hidden(sK3(X0,X1),X1)) & (r1_tarski(sK3(X0,X1),X0) | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [skolemisation 40,41]
47. r1_tarski(sK0,sK1) [cnf transformation 34] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. ~r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) [cnf transformation 34] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. r2_hidden(sK2(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 38] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
55. ~r2_hidden(sK2(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 38] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
56. r1_tarski(X3,X0) | ~r2_hidden(X3,X1) | k1_zfmisc_1(X0) != X1 [cnf transformation 42] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
57. r2_hidden(X3,X1) | ~r1_tarski(X3,X0) | k1_zfmisc_1(X0) != X1 [cnf transformation 42] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
62. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 31] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. r2_hidden(X3,k1_zfmisc_1(X0)) | ~r1_tarski(X3,X0) [equality resolution 57] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
68. ~r2_hidden(X3,k1_zfmisc_1(X0)) | r1_tarski(X3,X0) [equality resolution 56] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
70. 1 <=> r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) [avatar definition]
72. ~r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) <- (~1) [avatar component clause 70] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
73. ~1 [avatar split clause 48,70]
75. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
77. r1_tarski(sK0,sK1) <- (2) [avatar component clause 75] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
78. 2 [avatar split clause 47,75]
112. r1_tarski(sK2(k1_zfmisc_1(X4),X5),X4) | r1_tarski(k1_zfmisc_1(X4),X5) [resolution 54,68] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
120. ~r1_tarski(sK2(X1,k1_zfmisc_1(X2)),X2) | r1_tarski(X1,k1_zfmisc_1(X2)) [resolution 55,67] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
124. ~r1_tarski(X0,sK0) | r1_tarski(X0,sK1) <- (2) [resolution 62,77] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
138. r1_tarski(sK2(k1_zfmisc_1(sK0),X3),sK1) | r1_tarski(k1_zfmisc_1(sK0),X3) <- (2) [resolution 112,124] {a:2,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
141. r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) | r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) <- (2) [resolution 120,138] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
143. r1_tarski(k1_zfmisc_1(sK0),k1_zfmisc_1(sK1)) <- (2) [duplicate literal removal 141] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
145. $false <- (~1, 2) [subsumption resolution 143,72] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
146. 1 | ~2 [avatar contradiction clause 145]
147. $false [avatar sat refutation 73,78,146] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:36,thDist:-36}
% SZS output end Proof for zfmisc_1__t67_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 20
% Generated clauses: 65
% Active clauses: 29
% Passive clauses: 41
% Final active clauses: 28
% Final passive clauses: 11
% Main loop iterations started: 30
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 29
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 7
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

