i: [1,80,0,2,1,1,0,0] t37_relat_1
a: [3,123,0,2,1,80]
s: 123
i: [1,82,0,3,1,1,0,0] t37_relat_1
a: [3,118,0,3,1,82]
s: 118
i: [1,81,0,5,1,1,0,0] t37_relat_1
s: 81
i: [1,83,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,128,0,2,1,83]
s: 128
i: [1,84,0,2,1,0,0,1] fc1_xboole_0
a: [3,133,0,2,1,84]
s: 133
i: [1,108,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,143,0,2,1,108]
s: 143
i: [1,109,0,2,1,0,0,2] rc2_xboole_0
a: [3,148,0,2,1,109]
s: 148
i: [1,110,0,2,1,0,0,2] rc1_xboole_0
a: [3,153,0,2,1,110]
s: 153
i: [1,112,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,158,0,2,1,112]
s: 158
i: [1,111,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,163,0,2,1,111]
s: 163
i: [1,114,0,2,1,0,0,1] rc1_relat_1
a: [3,168,0,2,1,114]
s: 168
i: [1,113,0,2,1,0,0,1] rc1_relat_1
a: [3,173,0,2,1,113]
s: 173
i: [1,86,0,3,1,0,0,255] fc2_xboole_0
s: 86
i: [1,87,0,3,1,0,0,255] fc2_zfmisc_1
s: 87
i: [1,85,0,3,1,0,0,255] d2_xboole_0
a: [3,138,0,3,1,85]
s: 138
r: 128
i: [1,88,0,4,2,0,0,255] cc2_zfmisc_1
s: 88
d: [2,176,1,2,1,1,40,88,153]
a: [3,181,1,2,1,176]
s: 181
i: [1,89,0,4,2,0,0,255] cc1_zfmisc_1
s: 89
d: [2,177,1,2,1,1,40,88,133]
a: [3,186,1,2,1,177]
s: 186
i: [1,90,0,4,2,0,0,1] cc1_relat_1
s: 90
i: [1,101,0,4,1,0,0,255] existence_m1_subset_1
s: 101
d: [2,188,1,2,1,1,40,90,153]
a: [3,193,1,2,1,188]
s: 193
i: [1,102,0,4,1,0,0,255] fc3_xboole_0
s: 102
d: [2,189,1,2,1,1,40,90,133]
a: [3,198,1,2,1,189]
s: 198
i: [1,103,0,4,1,0,0,1] fc1_zfmisc_1
s: 103
i: [1,91,0,5,2,0,0,1] t6_boole
s: 91
i: [1,100,0,5,2,0,0,1] t7_boole
s: 100
i: [1,104,0,5,1,0,0,255] fc5_relat_1
s: 104
i: [1,92,0,6,2,0,0,255] dt_o_1_1_relat_1
s: 92
i: [1,94,0,6,2,0,0,1] d1_relat_1
s: 94
i: [1,97,0,6,2,0,0,255] t1_subset
s: 97
i: [1,106,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 106
i: [1,99,0,7,3,0,0,1] t8_boole
s: 99
d: [2,205,1,6,2,0,40,97,94]
s: 205
i: [1,105,0,7,1,0,0,255] commutativity_k2_tarski
s: 105
i: [1,95,0,8,2,0,0,1] d1_relat_1
s: 95
d: [2,206,1,6,2,0,40,106,94]
s: 206
i: [1,98,0,8,3,0,0,255] t2_subset
s: 98
d: [2,209,1,6,2,0,40,98,101]
s: 209
i: [1,107,0,8,1,0,0,255] fc7_relat_1
s: 107
i: [1,96,0,10,1,0,0,255] d5_tarski
s: 96
d: [2,212,2,6,2,0,40,209,106]
s: 212
i: [1,93,0,12,3,0,0,1] d1_relat_1
s: 93
d: [2,230,1,8,3,0,44,81,93]
s: 230
d: [2,200,1,7,3,0,44,91,91]
s: 200
d: [2,229,1,7,3,0,44,103,93]
s: 229
d: [2,232,2,8,3,0,40,230,209]
a: [3,239,2,6,2,232]
s: 239
d: [2,210,1,8,3,0,40,98,92]
s: 210
a: [3,236,2,2,1,232]
s: 236
d: [2,257,3,0,0,2,52,118,236,91]
e: 257
d: [2,258,3,0,0,3,52,133,118,236,99]
e: 258
d: [2,259,3,0,0,3,52,133,118,236,99]
e: 259
d: [2,243,3,4,2,1,40,239,209]
d: [2,244,3,2,1,2,21,243,123]
f: 257,244,232,80,82
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t37_relat_1
% SZS output start Proof for relat_1__t37_relat_1
1. ! [X0] : (v1_relat_1(X0) => (! [X1] : ! [X2] : ~r2_hidden(k4_tarski(X1,X2),X0) => k1_xboole_0 = X0)) [input t37_relat_1]
2. ~! [X0] : (v1_relat_1(X0) => (! [X1] : ! [X2] : ~r2_hidden(k4_tarski(X1,X2),X0) => k1_xboole_0 = X0)) [negated conjecture 1]
8. ! [X0] : (v1_relat_1(X0) <=> ! [X1] : ~(! [X2] : ! [X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) [input d1_relat_1]
18. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
34. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
35. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
38. ~! [X0] : (v1_relat_1(X0) => (! [X1,X2] : ~r2_hidden(k4_tarski(X1,X2),X0) => k1_xboole_0 = X0)) [flattening 2]
39. ! [X0] : (v1_relat_1(X0) <=> ! [X1] : ~(! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) [flattening 8]
42. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 34]
47. ? [X0] : ((k1_xboole_0 != X0 & ! [X1,X2] : ~r2_hidden(k4_tarski(X1,X2),X0)) & v1_relat_1(X0)) [ennf transformation 38]
48. ? [X0] : (k1_xboole_0 != X0 & ! [X1,X2] : ~r2_hidden(k4_tarski(X1,X2),X0) & v1_relat_1(X0)) [flattening 47]
52. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 35]
54. ! [X0] : (v1_relat_1(X0) <=> ! [X1] : (? [X2,X3] : k4_tarski(X2,X3) = X1 | ~r2_hidden(X1,X0))) [ennf transformation 39]
56. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 42]
57. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 56]
61. ? [X0] : (k1_xboole_0 != X0 & ! [X1,X2] : ~r2_hidden(k4_tarski(X1,X2),X0) & v1_relat_1(X0)) => (k1_xboole_0 != sK0 & ! [X2,X1] : ~r2_hidden(k4_tarski(X1,X2),sK0) & v1_relat_1(sK0)) [choice axiom]
62. k1_xboole_0 != sK0 & ! [X1,X2] : ~r2_hidden(k4_tarski(X1,X2),sK0) & v1_relat_1(sK0) [skolemisation 48,61]
63. ! [X0] : ((v1_relat_1(X0) | ? [X1] : (! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) & (! [X1] : (? [X2,X3] : k4_tarski(X2,X3) = X1 | ~r2_hidden(X1,X0)) | ~v1_relat_1(X0))) [nnf transformation 54]
64. ! [X0] : ((v1_relat_1(X0) | ? [X1] : (! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0))) & (! [X4] : (? [X5,X6] : k4_tarski(X5,X6) = X4 | ~r2_hidden(X4,X0)) | ~v1_relat_1(X0))) [rectify 63]
65. ! [X0] : (? [X1] : (! [X2,X3] : k4_tarski(X2,X3) != X1 & r2_hidden(X1,X0)) => (! [X3,X2] : k4_tarski(X2,X3) != sK1(X0) & r2_hidden(sK1(X0),X0))) [choice axiom]
66. ! [X4] : (? [X5,X6] : k4_tarski(X5,X6) = X4 => k4_tarski(sK2(X4),sK3(X4)) = X4) [choice axiom]
67. ! [X0] : ((v1_relat_1(X0) | (! [X2,X3] : k4_tarski(X2,X3) != sK1(X0) & r2_hidden(sK1(X0),X0))) & (! [X4] : (k4_tarski(sK2(X4),sK3(X4)) = X4 | ~r2_hidden(X4,X0)) | ~v1_relat_1(X0))) [skolemisation 64,66,65]
68. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK4(X0),X0)) [choice axiom]
69. ! [X0] : m1_subset_1(sK4(X0),X0) [skolemisation 18,68]
80. v1_relat_1(sK0) [cnf transformation 62] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
81. ~r2_hidden(k4_tarski(X1,X2),sK0) [cnf transformation 62] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
82. k1_xboole_0 != sK0 [cnf transformation 62] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
91. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 52] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
93. k4_tarski(sK2(X4),sK3(X4)) = X4 | ~r2_hidden(X4,X0) | ~v1_relat_1(X0) [cnf transformation 67] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
98. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
101. m1_subset_1(sK4(X0),X0) [cnf transformation 69] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. 1 <=> k1_xboole_0 = sK0 [avatar definition]
118. k1_xboole_0 != sK0 <- (~1) [avatar component clause 116] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
119. ~1 [avatar split clause 82,116]
121. 2 <=> v1_relat_1(sK0) [avatar definition]
123. v1_relat_1(sK0) <- (2) [avatar component clause 121] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
124. 2 [avatar split clause 80,121]
209. r2_hidden(sK4(X0),X0) | v1_xboole_0(X0) [resolution 98,101] {a:1,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
230. ~r2_hidden(X15,sK0) | ~r2_hidden(X15,X16) | ~v1_relat_1(X16) [superposition 81,93] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
232. ~r2_hidden(sK4(sK0),X1) | ~v1_relat_1(X1) | v1_xboole_0(sK0) [resolution 230,209] {a:2,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
234. 17 <=> v1_xboole_0(sK0) [avatar definition]
236. v1_xboole_0(sK0) <- (17) [avatar component clause 234] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
238. 18 <=> ! [X1] : (~r2_hidden(sK4(sK0),X1) | ~v1_relat_1(X1)) [avatar definition]
239. ~r2_hidden(sK4(sK0),X1) | ~v1_relat_1(X1) <- (18) [avatar component clause 238] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
240. 17 | 18 [avatar split clause 232,238,234]
243. ~v1_relat_1(sK0) | v1_xboole_0(sK0) <- (18) [resolution 239,209] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
244. v1_xboole_0(sK0) <- (2, 18) [subsumption resolution 243,123] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
245. 17 | ~2 | ~18 [avatar split clause 244,238,121,234]
257. $false <- (~1, 17) [unit resulting resolution 118,236,91] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
262. 1 | ~17 [avatar contradiction clause 257]
265. $false [avatar sat refutation 119,124,240,245,262] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:120,thDist:-120}
% SZS output end Proof for relat_1__t37_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 36
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 35
% Generated clauses: 107
% Active clauses: 47
% Passive clauses: 63
% Final active clauses: 46
% Final passive clauses: 14
% Main loop iterations started: 49
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 16
% 
% >>> Generating Inferences
% Binary resolution: 25
% Unit resulting resolution: 3
% Forward superposition: 6
% Backward superposition: 14
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 12
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

