i: [1,23,0,2,1,1,0,0] t42_orders_1
a: [3,50,0,2,1,23]
s: 50
i: [1,24,0,3,1,1,0,0] t42_orders_1
a: [3,45,0,3,1,24]
s: 45
i: [1,25,0,4,1,1,0,0] t42_orders_1
a: [3,40,0,4,1,25]
s: 40
i: [1,26,0,5,2,0,0,1] dt_k2_relat_1
s: 26
i: [1,27,0,7,2,0,0,255] involutiveness_k2_relat_1
s: 27
i: [1,32,0,8,3,0,0,1] d8_orders_1
s: 32
d: [2,53,1,5,2,1,40,32,45]
d: [2,54,1,3,1,2,21,53,50]
a: [3,58,1,3,1,54]
s: 58
i: [1,33,0,8,3,0,0,1] d8_orders_1
s: 33
d: [2,60,1,5,2,1,40,33,45]
d: [2,61,1,3,1,2,21,60,50]
a: [3,65,1,3,1,61]
s: 65
i: [1,34,0,8,3,0,0,1] d8_orders_1
s: 34
i: [1,35,0,8,3,0,0,1] d8_orders_1
s: 35
d: [2,67,1,5,2,1,40,34,45]
d: [2,68,1,3,1,2,21,67,50]
a: [3,72,1,3,1,68]
s: 72
i: [1,28,0,9,3,0,0,2] l58_orders_1
s: 28
d: [2,74,1,5,2,1,40,35,45]
d: [2,75,1,3,1,2,21,74,50]
a: [3,79,1,3,1,75]
s: 79
i: [1,29,0,9,3,0,0,2] l59_orders_1
s: 29
i: [1,30,0,9,3,0,0,2] l60_orders_1
s: 30
i: [1,31,0,9,3,0,0,2] l61_orders_1
s: 31
i: [1,36,0,17,6,0,0,1] d8_orders_1
s: 36
d: [2,81,1,12,4,0,44,28,27]
d: [2,82,1,9,3,0,21,81,26]
s: 82
d: [2,83,1,12,4,0,44,29,27]
d: [2,84,1,9,3,0,21,83,26]
s: 84
d: [2,85,1,12,4,0,44,30,27]
d: [2,86,1,9,3,0,21,85,26]
s: 86
d: [2,87,1,12,4,0,44,31,27]
d: [2,88,1,9,3,0,21,87,26]
s: 88
d: [2,90,1,24,7,0,40,36,31]
d: [2,91,1,21,6,0,21,90,26]
s: 91
d: [2,105,2,21,7,0,44,91,27]
d: [2,107,2,18,6,0,21,105,26]
s: 107
d: [2,104,2,22,7,0,40,91,30]
d: [2,106,2,20,6,0,18,104]
s: 106
d: [2,111,3,21,7,0,40,106,29]
d: [2,113,3,19,6,0,18,111]
s: 113
d: [2,115,4,20,7,0,40,113,28]
d: [2,117,4,18,6,0,18,115]
s: 117
d: [2,119,5,0,0,6,52,50,58,65,72,79,40,117]
e: 119
d: [2,120,5,14,5,1,40,117,40]
d: [2,127,5,11,4,2,21,120,79]
d: [2,128,5,9,3,3,21,127,50]
d: [2,129,5,6,2,4,21,128,72]
d: [2,130,5,3,1,5,21,129,65]
d: [2,131,5,0,0,6,21,130,58]
e: 131
f: 131,75,68,61,54,23,24,25
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t42_orders_1
% SZS output start Proof for orders_1__t42_orders_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r3_orders_1(X0,X1) => r3_orders_1(k2_relat_1(X0),X1))) [input t42_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (r3_orders_1(X0,X1) => r3_orders_1(k2_relat_1(X0),X1))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r3_orders_1(X0,X1) <=> (r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)))) [input d8_orders_1]
4. ! [X0] : (v1_relat_1(X0) => v1_relat_1(k2_relat_1(X0))) [input dt_k2_relat_1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_relat_2(X0,X1) => r1_relat_2(k2_relat_1(X0),X1))) [input l58_orders_1]
7. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r8_relat_2(X0,X1) => r8_relat_2(k2_relat_1(X0),X1))) [input l59_orders_1]
8. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r4_relat_2(X0,X1) => r4_relat_2(k2_relat_1(X0),X1))) [input l60_orders_1]
9. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r6_relat_2(X0,X1) => r6_relat_2(k2_relat_1(X0),X1))) [input l61_orders_1]
10. ? [X0] : (? [X1] : (~r3_orders_1(k2_relat_1(X0),X1) & r3_orders_1(X0,X1)) & v1_relat_1(X0)) [ennf transformation 2]
11. ! [X0] : (v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 4]
13. ! [X0] : (! [X1] : (r1_relat_2(k2_relat_1(X0),X1) | ~r1_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 6]
14. ! [X0] : (! [X1] : (r8_relat_2(k2_relat_1(X0),X1) | ~r8_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 7]
15. ! [X0] : (! [X1] : (r4_relat_2(k2_relat_1(X0),X1) | ~r4_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 8]
16. ! [X0] : (! [X1] : (r6_relat_2(k2_relat_1(X0),X1) | ~r6_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 9]
17. ! [X0] : (! [X1] : (r3_orders_1(X0,X1) <=> (r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 3]
18. ? [X0] : (? [X1] : (~r3_orders_1(k2_relat_1(X0),X1) & r3_orders_1(X0,X1)) & v1_relat_1(X0)) => (? [X1] : (~r3_orders_1(k2_relat_1(sK0),X1) & r3_orders_1(sK0,X1)) & v1_relat_1(sK0)) [choice axiom]
19. ? [X1] : (~r3_orders_1(k2_relat_1(sK0),X1) & r3_orders_1(sK0,X1)) => (~r3_orders_1(k2_relat_1(sK0),sK1) & r3_orders_1(sK0,sK1)) [choice axiom]
20. (~r3_orders_1(k2_relat_1(sK0),sK1) & r3_orders_1(sK0,sK1)) & v1_relat_1(sK0) [skolemisation 10,19,18]
21. ! [X0] : (! [X1] : ((r3_orders_1(X0,X1) | (~r6_relat_2(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1))) & ((r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r3_orders_1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 17]
22. ! [X0] : (! [X1] : ((r3_orders_1(X0,X1) | ~r6_relat_2(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1)) & ((r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r3_orders_1(X0,X1))) | ~v1_relat_1(X0)) [flattening 21]
23. v1_relat_1(sK0) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
24. r3_orders_1(sK0,sK1) [cnf transformation 20] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
25. ~r3_orders_1(k2_relat_1(sK0),sK1) [cnf transformation 20] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
26. v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 11] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
28. r1_relat_2(k2_relat_1(X0),X1) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 13] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
29. r8_relat_2(k2_relat_1(X0),X1) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 14] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
30. r4_relat_2(k2_relat_1(X0),X1) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 15] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
31. r6_relat_2(k2_relat_1(X0),X1) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 16] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
32. ~r3_orders_1(X0,X1) | r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. ~r3_orders_1(X0,X1) | r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. ~r3_orders_1(X0,X1) | r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
35. ~r3_orders_1(X0,X1) | r6_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
36. ~r6_relat_2(X0,X1) | r3_orders_1(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
38. 1 <=> r3_orders_1(k2_relat_1(sK0),sK1) [avatar definition]
40. ~r3_orders_1(k2_relat_1(sK0),sK1) <- (~1) [avatar component clause 38] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
41. ~1 [avatar split clause 25,38]
43. 2 <=> r3_orders_1(sK0,sK1) [avatar definition]
45. r3_orders_1(sK0,sK1) <- (2) [avatar component clause 43] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
46. 2 [avatar split clause 24,43]
48. 3 <=> v1_relat_1(sK0) [avatar definition]
50. v1_relat_1(sK0) <- (3) [avatar component clause 48] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
51. 3 [avatar split clause 23,48]
53. r1_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 32,45] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
54. r1_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 53,50] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
56. 4 <=> r1_relat_2(sK0,sK1) [avatar definition]
58. r1_relat_2(sK0,sK1) <- (4) [avatar component clause 56] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
59. 4 | ~2 | ~3 [avatar split clause 54,48,43,56]
60. r8_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 33,45] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
61. r8_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 60,50] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
63. 5 <=> r8_relat_2(sK0,sK1) [avatar definition]
65. r8_relat_2(sK0,sK1) <- (5) [avatar component clause 63] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66. 5 | ~2 | ~3 [avatar split clause 61,48,43,63]
67. r4_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 34,45] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
68. r4_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 67,50] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
70. 6 <=> r4_relat_2(sK0,sK1) [avatar definition]
72. r4_relat_2(sK0,sK1) <- (6) [avatar component clause 70] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
73. 6 | ~2 | ~3 [avatar split clause 68,48,43,70]
74. r6_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 35,45] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
75. r6_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 74,50] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
77. 7 <=> r6_relat_2(sK0,sK1) [avatar definition]
79. r6_relat_2(sK0,sK1) <- (7) [avatar component clause 77] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
80. 7 | ~2 | ~3 [avatar split clause 75,48,43,77]
90. r3_orders_1(k2_relat_1(X0),X1) | ~r4_relat_2(k2_relat_1(X0),X1) | ~r8_relat_2(k2_relat_1(X0),X1) | ~r1_relat_2(k2_relat_1(X0),X1) | ~v1_relat_1(k2_relat_1(X0)) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 36,31] {a:1,w:24,wCS:9600,sine:1,thAx:0,allAx:2,thDist:-2}
91. ~r4_relat_2(k2_relat_1(X0),X1) | r3_orders_1(k2_relat_1(X0),X1) | ~r8_relat_2(k2_relat_1(X0),X1) | ~r1_relat_2(k2_relat_1(X0),X1) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) [subsumption resolution 90,26] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
104. r3_orders_1(k2_relat_1(X0),X1) | ~r8_relat_2(k2_relat_1(X0),X1) | ~r1_relat_2(k2_relat_1(X0),X1) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 91,30] {a:2,w:22,wCS:8800,sine:1,thAx:0,allAx:3,thDist:-3}
106. ~r8_relat_2(k2_relat_1(X0),X1) | r3_orders_1(k2_relat_1(X0),X1) | ~r1_relat_2(k2_relat_1(X0),X1) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) [duplicate literal removal 104] {a:2,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
111. r3_orders_1(k2_relat_1(X0),X1) | ~r1_relat_2(k2_relat_1(X0),X1) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 106,29] {a:3,w:21,wCS:8400,sine:1,thAx:0,allAx:4,thDist:-4}
113. ~r1_relat_2(k2_relat_1(X0),X1) | r3_orders_1(k2_relat_1(X0),X1) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) [duplicate literal removal 111] {a:3,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
115. r3_orders_1(k2_relat_1(X0),X1) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 113,28] {a:4,w:20,wCS:8000,sine:1,thAx:0,allAx:5,thDist:-5}
117. r3_orders_1(k2_relat_1(X0),X1) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) [duplicate literal removal 115] {a:4,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
120. ~r6_relat_2(sK0,sK1) | ~v1_relat_1(sK0) | ~r4_relat_2(sK0,sK1) | ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1) [resolution 117,40] {a:5,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
127. ~v1_relat_1(sK0) | ~r4_relat_2(sK0,sK1) | ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1, 7) [subsumption resolution 120,79] {a:5,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
128. ~r4_relat_2(sK0,sK1) | ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1, 3, 7) [subsumption resolution 127,50] {a:5,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
129. ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1, 3, 6, 7) [subsumption resolution 128,72] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
130. ~r1_relat_2(sK0,sK1) <- (~1, 3, 5, 6, 7) [subsumption resolution 129,65] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
131. $false <- (~1, 3, 4, 5, 6, 7) [subsumption resolution 130,58] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
132. 1 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar contradiction clause 131]
134. $false [avatar sat refutation 41,46,51,59,66,73,80,132] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:75,thDist:-75}
% SZS output end Proof for orders_1__t42_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 14
% Generated clauses: 81
% Active clauses: 27
% Passive clauses: 29
% Final active clauses: 27
% Final passive clauses: 2
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Fw subsumption resolutions: 18
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 19
% Unit resulting resolution: 1
% Forward superposition: 13
% Backward superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

