i: [1,52,0,3,1,1,0,0] t88_xboole_1
a: [3,82,0,3,1,52]
s: 82
i: [1,53,0,7,1,1,0,0] t88_xboole_1
a: [3,77,0,7,1,53]
s: 77
i: [1,54,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,87,0,2,1,54]
s: 87
i: [1,55,0,2,1,0,0,2] fc1_xboole_0
a: [3,92,0,2,1,55]
s: 92
i: [1,72,0,2,1,0,0,2] rc2_xboole_0
a: [3,102,0,2,1,72]
s: 102
i: [1,73,0,2,1,0,0,2] rc1_xboole_0
a: [3,107,0,2,1,73]
s: 107
i: [1,66,0,3,1,0,0,255] reflexivity_r1_tarski
s: 66
i: [1,56,0,3,1,0,0,255] d2_xboole_0
a: [3,97,0,3,1,56]
s: 97
r: 87
i: [1,57,0,5,1,0,0,1] t4_boole
s: 57
i: [1,58,0,5,1,0,0,1] t3_boole
s: 58
i: [1,59,0,5,1,0,0,1] t1_boole
s: 59
i: [1,60,0,5,2,0,0,2] t6_boole
s: 60
i: [1,67,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 67
i: [1,69,0,6,2,0,0,1] fc4_xboole_0
s: 69
i: [1,70,0,6,2,0,0,1] fc5_xboole_0
s: 70
i: [1,71,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 71
d: [2,119,1,3,1,1,40,71,82]
a: [3,123,1,3,1,119]
s: 123
i: [1,64,0,7,3,0,0,2] t8_boole
s: 64
i: [1,68,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 68
i: [1,61,0,8,2,0,0,1] t83_xboole_1
s: 61
i: [1,62,0,8,2,0,0,1] t83_xboole_1
s: 62
d: [2,134,1,10,2,1,44,77,61]
a: [3,138,1,5,1,134]
s: 138
i: [1,63,0,8,2,0,0,255] l32_xboole_1
s: 63
d: [2,151,1,8,2,1,44,77,63]
a: [3,158,1,5,1,151]
s: 158
i: [1,65,0,13,1,0,0,1] t42_xboole_1
s: 65
d: [2,145,1,6,2,0,44,62,58]
d: [2,147,1,3,1,0,19,145]
s: 147
d: [2,144,1,6,2,0,44,62,57]
d: [2,148,1,3,1,0,19,144]
s: 148
d: [2,126,1,5,1,0,44,68,59]
s: 126
d: [2,184,2,3,1,0,40,147,71]
s: 184
d: [2,127,1,5,1,0,44,68,59]
s: 127
d: [2,128,1,5,1,0,44,59,68]
s: 128
d: [2,131,1,5,1,0,44,59,68]
s: 131
d: [2,149,1,6,2,0,44,63,58]
s: 149
d: [2,185,2,5,2,0,44,147,60]
s: 185
d: [2,152,1,6,2,0,44,58,63]
s: 152
d: [2,110,1,7,3,0,44,60,60]
s: 110
d: [2,198,3,2,1,1,40,185,138]
a: [3,202,3,2,1,198]
s: 202
d: [2,111,1,7,2,0,44,59,60]
s: 111
d: [2,187,2,5,2,0,44,148,60]
s: 187
d: [2,112,1,7,2,0,44,58,60]
s: 112
d: [2,113,1,7,2,0,44,57,60]
s: 113
d: [2,218,3,4,1,1,40,187,138]
a: [3,222,3,4,1,218]
s: 222
d: [2,150,1,9,3,0,44,63,61]
s: 150
d: [2,197,3,5,2,0,40,185,71]
s: 197
d: [2,153,1,9,3,0,44,62,63]
s: 153
d: [2,154,1,9,3,0,44,61,63]
s: 154
d: [2,240,2,6,2,0,40,150,66]
s: 240
d: [2,178,1,10,2,0,44,70,65]
s: 178
d: [2,188,2,7,2,0,44,126,60]
s: 188
d: [2,179,1,10,2,0,44,69,65]
s: 179
d: [2,164,1,11,1,0,44,65,57]
s: 164
d: [2,278,3,5,2,1,44,158,188]
d: [2,281,3,2,1,1,21,278,66]
a: [3,285,3,2,1,281]
s: 285
d: [2,170,1,14,2,0,44,65,63]
d: [2,183,1,12,2,0,22,170,59]
s: 183
d: [2,326,2,8,2,1,44,77,183]
d: [2,339,2,5,1,1,21,326,66]
a: [3,343,2,5,1,339]
s: 343
d: [2,348,3,0,0,2,52,82,343,61]
e: 348
d: [2,352,3,6,2,1,44,343,61]
d: [2,353,3,3,1,1,19,352]
d: [2,357,3,0,0,2,21,353,82]
e: 357
f: 357,339,52,53
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t88_xboole_1
% SZS output start Proof for xboole_1__t88_xboole_1
1. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) => k4_xboole_0(k2_xboole_0(X0,X1),X1) = X0) [input t88_xboole_1]
2. ~! [X0] : ! [X1] : (r1_xboole_0(X0,X1) => k4_xboole_0(k2_xboole_0(X0,X1),X1) = X0) [negated conjecture 1]
14. ! [X0] : ! [X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [input l32_xboole_1]
17. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
19. ! [X0] : k2_xboole_0(X0,k1_xboole_0) = X0 [input t1_boole]
21. ! [X0] : ! [X1] : ! [X2] : k4_xboole_0(k2_xboole_0(X0,X1),X2) = k2_xboole_0(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2)) [input t42_xboole_1]
25. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [input t83_xboole_1]
27. ~! [X0,X1] : (r1_xboole_0(X0,X1) => k4_xboole_0(k2_xboole_0(X0,X1),X1) = X0) [flattening 2]
28. ! [X0,X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [flattening 25]
29. ! [X0,X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [flattening 14]
32. ! [X0,X1,X2] : k4_xboole_0(k2_xboole_0(X0,X1),X2) = k2_xboole_0(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2)) [flattening 21]
33. ! [X0] : r1_tarski(X0,X0) [rectify 17]
35. ! [X0,X1] : (r1_tarski(X0,X1) => k1_xboole_0 = k4_xboole_0(X0,X1)) [unused predicate definition removal 29]
38. ? [X0,X1] : (k4_xboole_0(k2_xboole_0(X0,X1),X1) != X0 & r1_xboole_0(X0,X1)) [ennf transformation 27]
40. ! [X0,X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1)) [ennf transformation 35]
45. ? [X0,X1] : (k4_xboole_0(k2_xboole_0(X0,X1),X1) != X0 & r1_xboole_0(X0,X1)) => (sK0 != k4_xboole_0(k2_xboole_0(sK0,sK1),sK1) & r1_xboole_0(sK0,sK1)) [choice axiom]
46. sK0 != k4_xboole_0(k2_xboole_0(sK0,sK1),sK1) & r1_xboole_0(sK0,sK1) [skolemisation 38,45]
47. ! [X0,X1] : ((r1_xboole_0(X0,X1) | k4_xboole_0(X0,X1) != X0) & (k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1))) [nnf transformation 28]
52. r1_xboole_0(sK0,sK1) [cnf transformation 46] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. sK0 != k4_xboole_0(k2_xboole_0(sK0,sK1),sK1) [cnf transformation 46] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. k2_xboole_0(X0,k1_xboole_0) = X0 [cnf transformation 19] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1) [cnf transformation 47] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
65. k4_xboole_0(k2_xboole_0(X0,X1),X2) = k2_xboole_0(k4_xboole_0(X0,X2),k4_xboole_0(X1,X2)) [cnf transformation 32] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
66. r1_tarski(X0,X0) [cnf transformation 33] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
75. 1 <=> sK0 = k4_xboole_0(k2_xboole_0(sK0,sK1),sK1) [avatar definition]
77. sK0 != k4_xboole_0(k2_xboole_0(sK0,sK1),sK1) <- (~1) [avatar component clause 75] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
78. ~1 [avatar split clause 53,75]
80. 2 <=> r1_xboole_0(sK0,sK1) [avatar definition]
82. r1_xboole_0(sK0,sK1) <- (2) [avatar component clause 80] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. 2 [avatar split clause 52,80]
170. k4_xboole_0(k2_xboole_0(X6,X4),X5) = k2_xboole_0(k4_xboole_0(X6,X5),k1_xboole_0) | ~r1_tarski(X4,X5) [superposition 65,63] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
183. k4_xboole_0(X6,X5) = k4_xboole_0(k2_xboole_0(X6,X4),X5) | ~r1_tarski(X4,X5) [forward demodulation 170,59] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
326. sK0 != k4_xboole_0(sK0,sK1) | ~r1_tarski(sK1,sK1) <- (~1) [superposition 77,183] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
339. sK0 != k4_xboole_0(sK0,sK1) <- (~1) [subsumption resolution 326,66] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
341. 18 <=> sK0 = k4_xboole_0(sK0,sK1) [avatar definition]
343. sK0 != k4_xboole_0(sK0,sK1) <- (~18) [avatar component clause 341] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
344. ~18 | 1 [avatar split clause 339,75,341]
352. sK0 != sK0 | ~r1_xboole_0(sK0,sK1) <- (~18) [superposition 343,61] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
353. ~r1_xboole_0(sK0,sK1) <- (~18) [trivial inequality removal 352] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
357. $false <- (2, ~18) [subsumption resolution 353,82] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
358. ~2 | 18 [avatar contradiction clause 357]
359. $false [avatar sat refutation 78,83,344,358] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:101,thDist:-101}
% SZS output end Proof for xboole_1__t88_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 22
% Generated clauses: 230
% Active clauses: 45
% Passive clauses: 109
% Final active clauses: 44
% Final passive clauses: 55
% Main loop iterations started: 54
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Trivial inequalities: 7
% Fw subsumption resolutions: 7
% Fw demodulations: 14
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Equational tautologies: 3
% Forward subsumptions: 56
% Fw demodulations to eq. taut.: 4
% 
% >>> Generating Inferences
% Binary resolution: 14
% Unit resulting resolution: 1
% Forward superposition: 80
% Backward superposition: 66
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 7
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% forward subsumption: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

