i: [1,53,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,91,0,2,1,53]
s: 91
i: [1,54,0,2,1,0,0,255] fc1_xboole_0
a: [3,96,0,2,1,54]
s: 96
i: [1,76,0,2,1,0,0,2] rc2_xboole_0
a: [3,106,0,2,1,76]
s: 106
i: [1,77,0,2,1,0,0,2] rc1_xboole_0
a: [3,111,0,2,1,77]
s: 111
i: [1,52,0,9,1,1,0,0] t13_zfmisc_1
a: [3,86,0,9,1,52]
s: 86
i: [1,56,0,3,1,0,0,1] fc2_xboole_0
s: 56
i: [1,70,0,3,1,0,0,2] reflexivity_r1_tarski
s: 70
i: [1,55,0,3,1,0,0,255] d2_xboole_0
a: [3,101,0,3,1,55]
s: 101
r: 91
i: [1,71,0,4,1,0,0,1] fc3_xboole_0
s: 71
i: [1,57,0,5,1,0,0,255] t2_boole
s: 57
i: [1,58,0,5,2,0,0,255] t6_boole
s: 58
i: [1,63,0,5,2,0,0,2] t7_boole
s: 63
i: [1,72,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 72
i: [1,79,0,5,1,0,0,1] d2_tarski
s: 79
i: [1,81,0,5,1,0,0,1] d2_tarski
s: 81
i: [1,75,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 75
d: [2,119,1,5,1,0,40,75,79]
s: 119
i: [1,60,0,7,2,0,0,1] l1_zfmisc_1
s: 60
d: [2,121,1,4,1,0,40,60,70]
s: 121
i: [1,61,0,7,2,0,0,1] l1_zfmisc_1
s: 61
i: [1,62,0,7,3,0,0,2] t8_boole
s: 62
d: [2,122,2,4,1,0,40,121,75]
s: 122
i: [1,73,0,7,1,0,0,1] commutativity_k2_tarski
s: 73
d: [2,120,1,5,1,0,40,75,81]
s: 120
i: [1,74,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 74
i: [1,59,0,8,2,0,0,2] t28_xboole_1
s: 59
d: [2,142,1,11,2,1,44,86,59]
d: [2,146,1,6,1,1,19,142]
a: [3,150,1,6,1,146]
s: 150
d: [2,152,2,0,0,1,52,81,150,61]
e: 152
d: [2,153,2,5,1,1,40,150,61]
d: [2,155,2,0,0,1,21,153,81]
e: 155
f: 155,146,52
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t13_zfmisc_1
% SZS output start Proof for zfmisc_1__t13_zfmisc_1
1. ! [X0] : ! [X1] : k1_tarski(X0) = k3_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) [input t13_zfmisc_1]
2. ~! [X0] : ! [X1] : k1_tarski(X0) = k3_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) [negated conjecture 1]
6. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
17. ! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [input l1_zfmisc_1]
21. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
26. ~! [X0,X1] : k1_tarski(X0) = k3_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) [flattening 2]
27. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 21]
28. ! [X0,X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [flattening 17]
31. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 6]
34. ? [X0,X1] : k1_tarski(X0) != k3_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) [ennf transformation 26]
36. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 27]
40. ? [X0,X1] : k1_tarski(X0) != k3_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) => k1_tarski(sK0) != k3_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) [choice axiom]
41. k1_tarski(sK0) != k3_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) [skolemisation 34,40]
42. ! [X0,X1] : ((r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X1) | ~r1_tarski(k1_tarski(X0),X1))) [nnf transformation 28]
43. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 31]
44. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 43]
45. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 44]
46. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK2(X0,X1,X2) != X1 & sK2(X0,X1,X2) != X0) | ~r2_hidden(sK2(X0,X1,X2),X2)) & (sK2(X0,X1,X2) = X1 | sK2(X0,X1,X2) = X0 | r2_hidden(sK2(X0,X1,X2),X2)))) [choice axiom]
47. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK2(X0,X1,X2) != X1 & sK2(X0,X1,X2) != X0) | ~r2_hidden(sK2(X0,X1,X2),X2)) & (sK2(X0,X1,X2) = X1 | sK2(X0,X1,X2) = X0 | r2_hidden(sK2(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 45,46]
52. k1_tarski(sK0) != k3_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) [cnf transformation 41] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
61. r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 42] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. r2_hidden(X4,X2) | X0 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 47] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
80. r2_hidden(X4,X2) | k2_tarski(X4,X1) != X2 [equality resolution 65] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
81. r2_hidden(X4,k2_tarski(X4,X1)) [equality resolution 80] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
84. 1 <=> k1_tarski(sK0) = k3_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) [avatar definition]
86. k1_tarski(sK0) != k3_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) <- (~1) [avatar component clause 84] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. ~1 [avatar split clause 52,84]
142. k1_tarski(sK0) != k1_tarski(sK0) | ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) <- (~1) [superposition 86,59] {a:1,w:11,wCS:1100,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
146. ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) <- (~1) [trivial inequality removal 142] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
148. 7 <=> r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) [avatar definition]
150. ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) <- (~7) [avatar component clause 148] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
151. ~7 | 1 [avatar split clause 146,84,148]
153. ~r2_hidden(sK0,k2_tarski(sK0,sK1)) <- (~7) [resolution 150,61] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
155. $false <- (~7) [subsumption resolution 153,81] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
156. 7 [avatar contradiction clause 155]
157. $false [avatar sat refutation 87,151,156] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:30,thDist:-30}
% SZS output end Proof for zfmisc_1__t13_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 70
% Active clauses: 27
% Passive clauses: 45
% Final active clauses: 26
% Final passive clauses: 18
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 1
% Forward superposition: 7
% Backward superposition: 16
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 8
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 9338
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.004 s

