i: [1,75,0,2,1,1,0,0] t30_funct_4
a: [3,122,0,2,1,75]
s: 122
i: [1,74,0,2,1,1,0,0] t30_funct_4
a: [3,127,0,2,1,74]
s: 127
i: [1,73,0,2,1,1,0,0] t30_funct_4
a: [3,132,0,2,1,73]
s: 132
i: [1,72,0,2,1,1,0,0] t30_funct_4
a: [3,137,0,2,1,72]
s: 137
i: [1,77,0,10,2,1,0,0] t30_funct_4
a: [3,112,0,3,1,77]
s: 112
a: [3,115,0,7,1,77]
s: 115
i: [1,105,0,2,1,0,0,1] rc1_funct_1
a: [3,153,0,2,1,105]
s: 153
i: [1,104,0,2,1,0,0,1] rc1_funct_1
a: [3,158,0,2,1,104]
s: 158
i: [1,95,0,3,1,0,0,2] reflexivity_r1_tarski
s: 95
i: [1,107,0,3,1,0,0,1] d10_xboole_0
s: 107
i: [1,108,0,3,1,0,0,1] d10_xboole_0
s: 108
i: [1,94,0,4,1,0,0,255] existence_m1_subset_1
s: 94
i: [1,98,0,11,5,0,0,1] reflexivity_r1_partfun1
a: [3,141,0,4,2,98]
s: 141
d: [2,160,1,0,0,3,52,127,122,141]
e: 160
d: [2,161,1,0,0,3,52,137,132,141]
e: 161
d: [2,162,1,0,0,3,52,158,153,141]
e: 162
d: [2,163,1,2,1,2,40,141,122]
d: [2,169,1,0,0,3,21,163,127]
e: 169
d: [2,164,1,2,1,2,40,141,132]
d: [2,171,1,0,0,3,21,164,137]
e: 171
d: [2,165,1,2,1,2,40,141,153]
d: [2,173,1,0,0,3,21,165,158]
e: 173
r: 141
i: [1,96,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 96
i: [1,91,0,7,2,0,0,255] t3_subset
s: 91
i: [1,92,0,7,2,0,0,255] t3_subset
s: 92
d: [2,175,1,5,1,0,40,91,94]
s: 175
i: [1,97,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 97
a: [3,144,0,7,3,98]
s: 144
i: [1,78,0,8,3,0,0,255] cc2_relat_1
s: 78
i: [1,103,0,8,3,0,0,1] fc3_relat_1
s: 103
d: [2,177,1,6,2,0,40,78,94]
s: 177
i: [1,90,0,9,3,0,0,1] d10_xboole_0
s: 90
d: [2,178,1,7,3,0,40,78,92]
s: 178
i: [1,99,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,148,0,9,3,99]
s: 148
i: [1,79,0,10,4,0,0,255] cc3_funct_1
s: 79
d: [2,187,1,8,3,0,40,79,94]
s: 187
i: [1,93,0,11,3,0,0,1] t8_xboole_1
s: 93
d: [2,188,1,9,4,0,40,79,92]
s: 188
i: [1,100,0,12,5,0,0,1] dt_k1_funct_4
s: 100
i: [1,101,0,12,5,0,0,1] dt_k1_funct_4
s: 101
d: [2,201,1,12,5,1,44,101,115]
d: [2,204,1,10,4,2,21,201,122]
d: [2,205,1,8,3,3,21,204,127]
d: [2,206,1,6,2,4,21,205,132]
d: [2,207,1,4,1,5,21,206,137]
a: [3,211,1,4,1,207]
s: 211
i: [1,80,0,13,5,0,0,1] t25_funct_4
s: 80
r: 112
r: 115
a: [3,111,0,3,1,77]
s: 111
a: [3,116,0,7,1,77]
s: 116
i: [1,102,0,14,6,0,0,1] symmetry_r1_partfun1
s: 102
d: [2,229,1,11,5,1,40,102,111]
d: [2,231,1,9,4,2,21,229,122]
d: [2,232,1,7,3,3,21,231,127]
d: [2,233,1,5,2,4,21,232,132]
d: [2,234,1,3,1,5,21,233,137]
a: [3,238,1,3,1,234]
s: 238
i: [1,81,0,15,5,0,0,1] t29_funct_4
s: 81
d: [2,245,1,13,5,0,44,81,96]
d: [2,248,1,9,3,0,18,245]
s: 248
i: [1,84,0,15,6,0,0,1] t51_partfun1
s: 84
i: [1,85,0,15,6,0,0,1] t51_partfun1
s: 85
d: [2,252,2,10,4,0,40,248,178]
d: [2,256,2,8,3,0,18,252]
s: 256
i: [1,82,0,16,6,0,0,1] t28_funct_4
s: 82
d: [2,251,2,12,5,0,40,248,188]
d: [2,257,2,8,3,0,18,251]
s: 257
i: [1,83,0,16,6,0,0,1] t28_funct_4
s: 83
i: [1,86,0,18,6,0,0,1] t51_partfun1
s: 86
d: [2,184,2,10,2,0,40,90,175]
s: 184
i: [1,106,0,19,7,0,0,1] t51_partfun1
s: 106
d: [2,189,2,12,4,0,40,93,178]
s: 189
d: [2,279,0,15,6,0,21,106,103]
s: 279
d: [2,246,1,15,5,0,44,81,97]
s: 246
d: [2,197,2,14,5,0,40,188,93]
s: 197
d: [2,247,1,15,5,0,44,81,97]
s: 247
d: [2,299,3,21,9,0,40,197,279]
d: [2,304,3,19,8,0,21,299,178]
d: [2,305,3,17,7,0,21,304,178]
d: [2,306,3,15,6,0,21,305,188]
d: [2,307,3,13,5,0,21,306,188]
s: 307
d: [2,274,1,26,11,0,44,85,86]
d: [2,277,1,15,6,0,18,274]
s: 277
d: [2,286,1,15,6,0,44,279,97]
s: 286
d: [2,308,4,10,4,0,40,307,95]
s: 308
d: [2,287,1,15,6,0,44,279,97]
s: 287
d: [2,338,5,9,3,0,40,308,175]
s: 338
d: [2,190,1,16,4,0,40,93,90]
s: 190
d: [2,215,1,18,6,0,40,80,90]
s: 215
d: [2,339,5,13,5,0,40,308,248]
d: [2,345,5,9,3,0,18,339]
s: 345
d: [2,262,1,21,7,0,40,82,90]
s: 262
d: [2,352,6,21,7,0,40,338,102]
d: [2,353,6,17,5,0,18,352]
d: [2,354,6,13,4,0,21,353,187]
d: [2,355,6,9,3,0,21,354,177]
s: 355
d: [2,243,1,22,6,0,40,81,90]
s: 243
d: [2,269,1,31,7,0,40,83,93]
s: 269
d: [2,371,6,21,7,0,40,345,102]
d: [2,374,6,17,5,0,18,371]
d: [2,375,6,13,4,0,21,374,101]
d: [2,376,6,9,3,0,21,375,100]
s: 376
d: [2,325,2,15,6,0,44,277,97]
s: 325
d: [2,309,4,12,4,0,40,307,175]
s: 309
d: [2,326,2,15,6,0,44,277,97]
s: 326
d: [2,356,2,25,7,0,40,190,81]
d: [2,366,2,20,6,0,21,356,80]
s: 366
d: [2,433,3,20,6,1,44,116,366]
d: [2,436,3,13,5,1,19,433]
d: [2,465,3,11,4,2,21,436,122]
d: [2,466,3,9,3,3,21,465,127]
d: [2,467,3,7,2,4,21,466,132]
d: [2,468,3,5,1,5,21,467,137]
a: [3,472,3,5,1,468]
s: 472
d: [2,474,4,0,0,6,52,137,132,127,122,111,472,82]
e: 474
d: [2,475,4,11,5,1,40,472,82]
d: [2,478,4,8,4,2,21,475,111]
d: [2,479,4,6,3,3,21,478,122]
d: [2,480,4,4,2,4,21,479,127]
d: [2,481,4,2,1,5,21,480,132]
d: [2,482,4,0,0,6,21,481,137]
e: 482
r: 111
r: 116
i: [1,76,0,10,2,1,0,0] t30_funct_4
d: [2,484,4,9,4,2,21,475,122]
d: [2,485,4,7,3,3,21,484,127]
d: [2,486,4,5,2,4,21,485,132]
d: [2,487,4,3,1,5,21,486,137]
d: [2,284,2,11,5,1,40,279,211]
d: [2,489,2,9,4,2,21,284,122]
d: [2,490,2,7,3,3,21,489,127]
d: [2,491,2,5,2,4,21,490,132]
d: [2,492,2,3,1,5,21,491,137]
f: 492,487,468,207,72,73,74,75,76,77
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t30_funct_4
% SZS output start Proof for funct_4__t30_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) <=> k2_xboole_0(X0,X1) = k1_funct_4(X0,X1)))) [input t30_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) <=> k2_xboole_0(X0,X1) = k1_funct_4(X0,X1)))) [negated conjecture 1]
6. ! [X0] : ! [X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [input d10_xboole_0]
7. ! [X0,X1] : ((v1_funct_1(X1) & v1_relat_1(X1) & v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k1_funct_4(X0,X1)) & v1_relat_1(k1_funct_4(X0,X1)))) [input dt_k1_funct_4]
12. ! [X0,X1] : ((v1_relat_1(X1) & v1_relat_1(X0)) => v1_relat_1(k2_xboole_0(X0,X1))) [input fc3_relat_1]
19. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => r1_tarski(X0,k1_funct_4(X1,X0)))) [input t25_funct_4]
20. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) <=> r1_tarski(X0,k1_funct_4(X0,X1))))) [input t28_funct_4]
21. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => r1_tarski(k1_funct_4(X0,X1),k2_xboole_0(X0,X1)))) [input t29_funct_4]
23. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_partfun1(X0,X1) <=> ? [X2] : (k2_xboole_0(X0,X1) = X2 & v1_funct_1(X2) & v1_relat_1(X2))))) [input t51_partfun1]
24. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [input t8_xboole_1]
25. ! [X0,X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [flattening 6]
27. ! [X0,X1,X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [flattening 24]
30. ? [X0] : (? [X1] : ((r1_partfun1(X0,X1) <~> k2_xboole_0(X0,X1) = k1_funct_4(X0,X1)) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
31. ? [X0] : (? [X1] : ((r1_partfun1(X0,X1) <~> k2_xboole_0(X0,X1) = k1_funct_4(X0,X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 30]
35. ! [X0] : (! [X1] : (r1_tarski(X0,k1_funct_4(X1,X0)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 19]
36. ! [X0] : (! [X1] : (r1_tarski(X0,k1_funct_4(X1,X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 35]
37. ! [X0] : (! [X1] : (r1_tarski(k1_funct_4(X0,X1),k2_xboole_0(X0,X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 21]
38. ! [X0] : (! [X1] : (r1_tarski(k1_funct_4(X0,X1),k2_xboole_0(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 37]
39. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> r1_tarski(X0,k1_funct_4(X0,X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 20]
40. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> r1_tarski(X0,k1_funct_4(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 39]
41. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> ? [X2] : (k2_xboole_0(X0,X1) = X2 & v1_funct_1(X2) & v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 23]
42. ! [X0] : (! [X1] : ((r1_partfun1(X0,X1) <=> ? [X2] : (k2_xboole_0(X0,X1) = X2 & v1_funct_1(X2) & v1_relat_1(X2))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 41]
43. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | (~r1_tarski(X2,X1) | ~r1_tarski(X0,X1))) [ennf transformation 27]
44. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1)) [flattening 43]
49. ! [X0,X1] : ((v1_funct_1(k1_funct_4(X0,X1)) & v1_relat_1(k1_funct_4(X0,X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 7]
50. ! [X0,X1] : ((v1_funct_1(k1_funct_4(X0,X1)) & v1_relat_1(k1_funct_4(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 49]
53. ! [X0,X1] : (v1_relat_1(k2_xboole_0(X0,X1)) | (~v1_relat_1(X1) | ~v1_relat_1(X0))) [ennf transformation 12]
54. ! [X0,X1] : (v1_relat_1(k2_xboole_0(X0,X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0)) [flattening 53]
55. ? [X0] : (? [X1] : (((k2_xboole_0(X0,X1) != k1_funct_4(X0,X1) | ~r1_partfun1(X0,X1)) & (k2_xboole_0(X0,X1) = k1_funct_4(X0,X1) | r1_partfun1(X0,X1))) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [nnf transformation 31]
56. ? [X0] : (? [X1] : ((k2_xboole_0(X0,X1) != k1_funct_4(X0,X1) | ~r1_partfun1(X0,X1)) & (k2_xboole_0(X0,X1) = k1_funct_4(X0,X1) | r1_partfun1(X0,X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 55]
57. ? [X0] : (? [X1] : ((k2_xboole_0(X0,X1) != k1_funct_4(X0,X1) | ~r1_partfun1(X0,X1)) & (k2_xboole_0(X0,X1) = k1_funct_4(X0,X1) | r1_partfun1(X0,X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : ((k2_xboole_0(sK0,X1) != k1_funct_4(sK0,X1) | ~r1_partfun1(sK0,X1)) & (k2_xboole_0(sK0,X1) = k1_funct_4(sK0,X1) | r1_partfun1(sK0,X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
58. ? [X1] : ((k2_xboole_0(sK0,X1) != k1_funct_4(sK0,X1) | ~r1_partfun1(sK0,X1)) & (k2_xboole_0(sK0,X1) = k1_funct_4(sK0,X1) | r1_partfun1(sK0,X1)) & v1_funct_1(X1) & v1_relat_1(X1)) => ((k2_xboole_0(sK0,sK1) != k1_funct_4(sK0,sK1) | ~r1_partfun1(sK0,sK1)) & (k2_xboole_0(sK0,sK1) = k1_funct_4(sK0,sK1) | r1_partfun1(sK0,sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
59. ((k2_xboole_0(sK0,sK1) != k1_funct_4(sK0,sK1) | ~r1_partfun1(sK0,sK1)) & (k2_xboole_0(sK0,sK1) = k1_funct_4(sK0,sK1) | r1_partfun1(sK0,sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 56,58,57]
60. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ~r1_tarski(X0,k1_funct_4(X0,X1))) & (r1_tarski(X0,k1_funct_4(X0,X1)) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 40]
61. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ! [X2] : (k2_xboole_0(X0,X1) != X2 | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X2] : (k2_xboole_0(X0,X1) = X2 & v1_funct_1(X2) & v1_relat_1(X2)) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 42]
62. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ! [X2] : (k2_xboole_0(X0,X1) != X2 | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X3] : (k2_xboole_0(X0,X1) = X3 & v1_funct_1(X3) & v1_relat_1(X3)) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [rectify 61]
63. ! [X1,X0] : (? [X3] : (k2_xboole_0(X0,X1) = X3 & v1_funct_1(X3) & v1_relat_1(X3)) => (k2_xboole_0(X0,X1) = sK2(X0,X1) & v1_funct_1(sK2(X0,X1)) & v1_relat_1(sK2(X0,X1)))) [choice axiom]
64. ! [X0] : (! [X1] : (((r1_partfun1(X0,X1) | ! [X2] : (k2_xboole_0(X0,X1) != X2 | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & ((k2_xboole_0(X0,X1) = sK2(X0,X1) & v1_funct_1(sK2(X0,X1)) & v1_relat_1(sK2(X0,X1))) | ~r1_partfun1(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [skolemisation 62,63]
65. ! [X0,X1] : ((X0 = X1 | (~r1_tarski(X1,X0) | ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [nnf transformation 25]
66. ! [X0,X1] : ((X0 = X1 | ~r1_tarski(X1,X0) | ~r1_tarski(X0,X1)) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [flattening 65]
72. v1_relat_1(sK0) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
73. v1_funct_1(sK0) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
74. v1_relat_1(sK1) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
75. v1_funct_1(sK1) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
76. k2_xboole_0(sK0,sK1) = k1_funct_4(sK0,sK1) | r1_partfun1(sK0,sK1) [cnf transformation 59] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
77. k2_xboole_0(sK0,sK1) != k1_funct_4(sK0,sK1) | ~r1_partfun1(sK0,sK1) [cnf transformation 59] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
80. r1_tarski(X0,k1_funct_4(X1,X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
81. r1_tarski(k1_funct_4(X0,X1),k2_xboole_0(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 38] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
82. r1_tarski(X0,k1_funct_4(X0,X1)) | ~r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 60] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. r1_partfun1(X0,X1) | k2_xboole_0(X0,X1) != X2 | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 64] {a:0,w:20,wCS:8000,sine:1,thAx:0,allAx:2,thDist:-2}
90. ~r1_tarski(X1,X0) | X0 = X1 | ~r1_tarski(X0,X1) [cnf transformation 66] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
93. r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1) [cnf transformation 44] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. v1_funct_1(k1_funct_4(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 50] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
103. v1_relat_1(k2_xboole_0(X0,X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 54] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
106. r1_partfun1(X0,X1) | ~v1_funct_1(k2_xboole_0(X0,X1)) | ~v1_relat_1(k2_xboole_0(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [equality resolution 87] {a:0,w:19,wCS:7600,sine:1,thAx:0,allAx:2,thDist:-2}
110. 1 <=> r1_partfun1(sK0,sK1) [avatar definition]
114. 2 <=> k2_xboole_0(sK0,sK1) = k1_funct_4(sK0,sK1) [avatar definition]
115. k2_xboole_0(sK0,sK1) = k1_funct_4(sK0,sK1) <- (2) [avatar component clause 114] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. k2_xboole_0(sK0,sK1) != k1_funct_4(sK0,sK1) <- (~2) [avatar component clause 114] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
117. ~1 | ~2 [avatar split clause 77,114,110]
118. 1 | 2 [avatar split clause 76,114,110]
120. 3 <=> v1_funct_1(sK1) [avatar definition]
122. v1_funct_1(sK1) <- (3) [avatar component clause 120] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
123. 3 [avatar split clause 75,120]
125. 4 <=> v1_relat_1(sK1) [avatar definition]
127. v1_relat_1(sK1) <- (4) [avatar component clause 125] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. 4 [avatar split clause 74,125]
130. 5 <=> v1_funct_1(sK0) [avatar definition]
132. v1_funct_1(sK0) <- (5) [avatar component clause 130] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
133. 5 [avatar split clause 73,130]
135. 6 <=> v1_relat_1(sK0) [avatar definition]
137. v1_relat_1(sK0) <- (6) [avatar component clause 135] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
138. 6 [avatar split clause 72,135]
190. ~r1_tarski(X4,k2_xboole_0(X5,X3)) | ~r1_tarski(X5,X4) | k2_xboole_0(X5,X3) = X4 | ~r1_tarski(X3,X4) [resolution 93,90] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
201. v1_funct_1(k2_xboole_0(sK0,sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2) [superposition 101,115] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
204. v1_funct_1(k2_xboole_0(sK0,sK1)) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3) [subsumption resolution 201,122] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
205. v1_funct_1(k2_xboole_0(sK0,sK1)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3, 4) [subsumption resolution 204,127] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
206. v1_funct_1(k2_xboole_0(sK0,sK1)) | ~v1_relat_1(sK0) <- (2, 3, 4, 5) [subsumption resolution 205,132] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
207. v1_funct_1(k2_xboole_0(sK0,sK1)) <- (2, 3, 4, 5, 6) [subsumption resolution 206,137] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
209. 12 <=> v1_funct_1(k2_xboole_0(sK0,sK1)) [avatar definition]
211. v1_funct_1(k2_xboole_0(sK0,sK1)) <- (12) [avatar component clause 209] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
212. 12 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar split clause 207,135,130,125,120,114,209]
279. ~v1_funct_1(k2_xboole_0(X0,X1)) | r1_partfun1(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 106,103] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
284. r1_partfun1(sK0,sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (12) [resolution 279,211] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
356. ~r1_tarski(X0,k1_funct_4(X0,X1)) | k2_xboole_0(X0,X1) = k1_funct_4(X0,X1) | ~r1_tarski(X1,k1_funct_4(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [resolution 190,81] {a:2,w:25,wCS:10000,sine:1,thAx:0,allAx:3,thDist:-3}
366. k2_xboole_0(X0,X1) = k1_funct_4(X0,X1) | ~r1_tarski(X0,k1_funct_4(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [subsumption resolution 356,80] {a:2,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
433. k2_xboole_0(sK0,sK1) != k2_xboole_0(sK0,sK1) | ~r1_tarski(sK0,k1_funct_4(sK0,sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~2) [superposition 116,366] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
436. ~r1_tarski(sK0,k1_funct_4(sK0,sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~2) [trivial inequality removal 433] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
465. ~r1_tarski(sK0,k1_funct_4(sK0,sK1)) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~2, 3) [subsumption resolution 436,122] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
466. ~r1_tarski(sK0,k1_funct_4(sK0,sK1)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~2, 3, 4) [subsumption resolution 465,127] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
467. ~r1_tarski(sK0,k1_funct_4(sK0,sK1)) | ~v1_relat_1(sK0) <- (~2, 3, 4, 5) [subsumption resolution 466,132] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
468. ~r1_tarski(sK0,k1_funct_4(sK0,sK1)) <- (~2, 3, 4, 5, 6) [subsumption resolution 467,137] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
470. 15 <=> r1_tarski(sK0,k1_funct_4(sK0,sK1)) [avatar definition]
472. ~r1_tarski(sK0,k1_funct_4(sK0,sK1)) <- (~15) [avatar component clause 470] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
473. ~15 | 2 | ~3 | ~4 | ~5 | ~6 [avatar split clause 468,135,130,125,120,114,470]
475. ~r1_partfun1(sK0,sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~15) [resolution 472,82] {a:4,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
484. ~r1_partfun1(sK0,sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3, ~15) [subsumption resolution 475,122] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
485. ~r1_partfun1(sK0,sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3, 4, ~15) [subsumption resolution 484,127] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
486. ~r1_partfun1(sK0,sK1) | ~v1_relat_1(sK0) <- (3, 4, 5, ~15) [subsumption resolution 485,132] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
487. ~r1_partfun1(sK0,sK1) <- (3, 4, 5, 6, ~15) [subsumption resolution 486,137] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
488. ~1 | ~3 | ~4 | ~5 | ~6 | 15 [avatar split clause 487,470,135,130,125,120,110]
489. r1_partfun1(sK0,sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3, 12) [subsumption resolution 284,122] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
490. r1_partfun1(sK0,sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3, 4, 12) [subsumption resolution 489,127] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
491. r1_partfun1(sK0,sK1) | ~v1_relat_1(sK0) <- (3, 4, 5, 12) [subsumption resolution 490,132] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
492. r1_partfun1(sK0,sK1) <- (3, 4, 5, 6, 12) [subsumption resolution 491,137] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
493. 1 | ~3 | ~4 | ~5 | ~6 | ~12 [avatar split clause 492,209,135,130,125,120,110]
499. $false [avatar sat refutation 117,118,123,128,133,138,212,473,488,493] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:175,thDist:-175}
% SZS output end Proof for funct_4__t30_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 34
% Generated clauses: 361
% Active clauses: 66
% Passive clauses: 103
% Final active clauses: 61
% Final passive clauses: 31
% Main loop iterations started: 72
% 
% >>> Simplifying Inferences
% Duplicate literals: 203
% Trivial inequalities: 1
% Fw subsumption resolutions: 71
% Fw demodulations: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 20
% Equational tautologies: 2
% Forward subsumptions: 74
% 
% >>> Generating Inferences
% Binary resolution: 81
% Unit resulting resolution: 4
% Forward superposition: 58
% Backward superposition: 19
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 27
% SAT solver unit clauses: 6
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 5117
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% preprocessing: 0.002 s ( own 0.001 s ) 
% unit resulting resolution: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.005 s ( own 0.001 s ) 

