i: [1,89,0,2,1,1,0,0] t167_relat_1
a: [3,144,0,2,1,89]
s: 144
i: [1,92,0,3,1,1,0,0] t167_relat_1
a: [3,129,0,3,1,92]
s: 129
i: [1,90,0,3,1,1,0,0] t167_relat_1
a: [3,139,0,3,1,90]
s: 139
i: [1,91,0,4,1,1,0,0] t167_relat_1
a: [3,134,0,4,1,91]
s: 134
i: [1,93,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,149,0,2,1,93]
s: 149
i: [1,94,0,2,1,0,0,255] fc1_xboole_0
a: [3,154,0,2,1,94]
s: 154
i: [1,119,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,164,0,2,1,119]
s: 164
i: [1,120,0,2,1,0,0,2] rc2_xboole_0
a: [3,169,0,2,1,120]
s: 169
i: [1,121,0,2,1,0,0,2] rc1_xboole_0
a: [3,174,0,2,1,121]
s: 174
i: [1,123,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,179,0,2,1,123]
s: 179
i: [1,122,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,184,0,2,1,122]
s: 184
i: [1,125,0,2,1,0,0,1] rc1_relat_1
a: [3,189,0,2,1,125]
s: 189
i: [1,124,0,2,1,0,0,1] rc1_relat_1
a: [3,194,0,2,1,124]
s: 194
i: [1,117,0,3,1,0,0,2] reflexivity_r1_tarski
s: 117
i: [1,95,0,3,1,0,0,255] d2_xboole_0
a: [3,159,0,3,1,95]
s: 159
r: 149
i: [1,96,0,4,2,0,0,255] cc2_zfmisc_1
s: 96
d: [2,197,1,2,1,1,40,96,174]
a: [3,202,1,2,1,197]
s: 202
i: [1,98,0,4,2,0,0,255] cc1_zfmisc_1
s: 98
d: [2,198,1,2,1,1,40,96,154]
a: [3,207,1,2,1,198]
s: 207
i: [1,99,0,4,2,0,0,1] cc1_relat_1
s: 99
i: [1,116,0,4,1,0,0,2] existence_m1_subset_1
s: 116
d: [2,209,1,2,1,1,40,99,174]
a: [3,214,1,2,1,209]
s: 214
i: [1,100,0,5,2,0,0,255] t6_boole
s: 100
d: [2,210,1,2,1,1,40,99,154]
a: [3,219,1,2,1,210]
s: 219
i: [1,101,0,5,2,0,0,1] fc11_relat_1
s: 101
i: [1,102,0,5,2,0,0,1] fc5_xtuple_0
s: 102
i: [1,113,0,5,2,0,0,1] t7_boole
s: 113
i: [1,108,0,6,2,0,0,1] t1_subset
s: 108
d: [2,227,1,3,1,1,40,113,134]
a: [3,231,1,3,1,227]
s: 231
i: [1,118,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 118
i: [1,103,0,7,3,0,0,1] fc9_relat_1
s: 103
d: [2,239,2,2,1,1,40,231,101]
a: [3,243,2,2,1,239]
s: 243
i: [1,104,0,7,3,0,0,255] fc25_relat_1
s: 104
d: [2,233,1,4,1,1,40,108,134]
a: [3,237,1,4,1,233]
s: 237
i: [1,110,0,7,2,0,0,2] t3_subset
s: 110
i: [1,111,0,7,2,0,0,2] t3_subset
s: 111
d: [2,245,1,4,1,1,40,118,134]
a: [3,249,1,4,1,245]
s: 249
i: [1,112,0,7,3,0,0,255] t8_boole
s: 112
d: [2,225,1,5,2,0,40,101,99]
s: 225
i: [1,97,0,8,3,0,0,3] cc2_relat_1
s: 97
i: [1,109,0,8,3,0,0,1] t2_subset
s: 109
d: [2,226,1,5,2,0,40,101,96]
s: 226
i: [1,106,0,9,3,0,0,1] d19_relat_1
s: 106
d: [2,252,1,5,1,0,40,110,116]
s: 252
i: [1,107,0,9,3,0,0,1] d19_relat_1
s: 107
i: [1,115,0,9,3,0,0,3] t5_subset
s: 115
d: [2,254,1,6,2,0,40,97,116]
s: 254
i: [1,114,0,10,3,0,0,3] t4_subset
s: 114
d: [2,256,1,6,2,0,40,109,116]
s: 256
d: [2,260,1,6,2,0,40,107,117]
s: 260
d: [2,221,1,7,3,0,44,100,100]
s: 221
d: [2,266,2,6,2,0,40,256,118]
s: 266
d: [2,255,1,7,3,0,40,97,111]
s: 255
d: [2,262,1,7,2,0,40,115,116]
s: 262
d: [2,263,1,8,3,0,40,115,111]
s: 263
d: [2,264,1,8,2,0,40,114,116]
s: 264
d: [2,272,2,6,2,0,40,262,256]
s: 272
d: [2,265,1,9,3,0,40,114,111]
s: 265
d: [2,277,3,6,2,0,40,272,99]
s: 277
d: [2,258,1,10,3,0,40,109,111]
s: 258
d: [2,270,2,10,4,0,40,255,106]
s: 270
d: [2,285,3,7,3,1,40,270,139]
d: [2,288,3,5,2,2,21,285,144]
a: [3,292,3,2,1,288]
s: 292
d: [2,276,2,10,2,0,40,264,256]
s: 276
d: [2,278,3,6,2,0,40,272,96]
s: 278
d: [2,282,2,10,3,0,40,258,118]
s: 282
d: [2,274,2,11,4,0,40,263,106]
s: 274
d: [2,305,3,7,3,1,40,274,134]
d: [2,307,3,5,2,2,21,305,144]
s: 307
d: [2,280,2,12,4,0,40,265,106]
s: 280
d: [2,308,4,2,1,3,40,307,139]
a: [3,313,4,2,1,308]
s: 313
d: [2,315,3,8,3,1,40,280,134]
d: [2,317,3,6,2,2,21,315,144]
s: 317
d: [2,302,3,12,3,0,40,276,109]
d: [2,303,3,10,2,0,21,302,272]
s: 303
d: [2,318,4,3,1,3,40,317,139]
a: [3,323,4,3,1,318]
s: 323
d: [2,332,5,0,0,3,52,313,129,323,109]
e: 332
d: [2,333,5,5,2,1,40,323,109]
d: [2,335,5,3,1,2,21,333,313]
d: [2,336,5,0,0,3,21,335,129]
e: 336
f: 336,318,308,89,90,91,92
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t167_relat_1
% SZS output start Proof for relat_1__t167_relat_1
1. ! [X0] : ! [X1] : ((v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (r2_hidden(X2,k10_xtuple_0(X1)) => r2_hidden(X2,X0))) [input t167_relat_1]
2. ~! [X0] : ! [X1] : ((v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (r2_hidden(X2,k10_xtuple_0(X1)) => r2_hidden(X2,X0))) [negated conjecture 1]
9. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v5_relat_1(X1,X0) <=> r1_tarski(k10_xtuple_0(X1),X0))) [input d19_relat_1]
29. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
30. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
31. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
32. ! [X0] : ! [X1] : ! [X2] : ~(v1_xboole_0(X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) [input t5_subset]
36. ~! [X0,X1] : ((v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (r2_hidden(X2,k10_xtuple_0(X1)) => r2_hidden(X2,X0))) [flattening 2]
37. ! [X0,X1] : (v1_relat_1(X1) => (v5_relat_1(X1,X0) <=> r1_tarski(k10_xtuple_0(X1),X0))) [flattening 9]
39. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 29]
40. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 30]
43. ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 31]
44. ! [X0,X1,X2] : ~(v1_xboole_0(X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) [flattening 32]
47. ? [X0,X1] : (? [X2] : (~r2_hidden(X2,X0) & r2_hidden(X2,k10_xtuple_0(X1))) & (v5_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 36]
48. ? [X0,X1] : (? [X2] : (~r2_hidden(X2,X0) & r2_hidden(X2,k10_xtuple_0(X1))) & v5_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 47]
62. ! [X0,X1] : ((v5_relat_1(X1,X0) <=> r1_tarski(k10_xtuple_0(X1),X0)) | ~v1_relat_1(X1)) [ennf transformation 37]
64. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 39]
65. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 64]
68. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))) [ennf transformation 43]
69. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [flattening 68]
70. ! [X0,X1,X2] : (~v1_xboole_0(X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [ennf transformation 44]
72. ? [X0,X1] : (? [X2] : (~r2_hidden(X2,X0) & r2_hidden(X2,k10_xtuple_0(X1))) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (~r2_hidden(X2,sK0) & r2_hidden(X2,k10_xtuple_0(sK1))) & v5_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
73. ? [X2] : (~r2_hidden(X2,sK0) & r2_hidden(X2,k10_xtuple_0(sK1))) => (~r2_hidden(sK2,sK0) & r2_hidden(sK2,k10_xtuple_0(sK1))) [choice axiom]
74. (~r2_hidden(sK2,sK0) & r2_hidden(sK2,k10_xtuple_0(sK1))) & v5_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 48,73,72]
75. ! [X0,X1] : (((v5_relat_1(X1,X0) | ~r1_tarski(k10_xtuple_0(X1),X0)) & (r1_tarski(k10_xtuple_0(X1),X0) | ~v5_relat_1(X1,X0))) | ~v1_relat_1(X1)) [nnf transformation 62]
76. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 40]
89. v1_relat_1(sK1) [cnf transformation 74] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
90. v5_relat_1(sK1,sK0) [cnf transformation 74] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
91. r2_hidden(sK2,k10_xtuple_0(sK1)) [cnf transformation 74] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
92. ~r2_hidden(sK2,sK0) [cnf transformation 74] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
106. r1_tarski(k10_xtuple_0(X1),X0) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 75] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
109. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 65] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
111. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 76] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
114. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1) [cnf transformation 69] {a:0,w:10,wCS:4000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
115. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~v1_xboole_0(X2) | ~r2_hidden(X0,X1) [cnf transformation 70] {a:0,w:9,wCS:3600,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
127. 1 <=> r2_hidden(sK2,sK0) [avatar definition]
129. ~r2_hidden(sK2,sK0) <- (~1) [avatar component clause 127] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
130. ~1 [avatar split clause 92,127]
132. 2 <=> r2_hidden(sK2,k10_xtuple_0(sK1)) [avatar definition]
134. r2_hidden(sK2,k10_xtuple_0(sK1)) <- (2) [avatar component clause 132] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
135. 2 [avatar split clause 91,132]
137. 3 <=> v5_relat_1(sK1,sK0) [avatar definition]
139. v5_relat_1(sK1,sK0) <- (3) [avatar component clause 137] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
140. 3 [avatar split clause 90,137]
142. 4 <=> v1_relat_1(sK1) [avatar definition]
144. v1_relat_1(sK1) <- (4) [avatar component clause 142] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
145. 4 [avatar split clause 89,142]
263. ~r1_tarski(X4,X2) | ~r2_hidden(X3,X4) | ~v1_xboole_0(X2) [resolution 115,111] {a:1,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
265. ~r1_tarski(X4,X3) | ~r2_hidden(X2,X4) | m1_subset_1(X2,X3) [resolution 114,111] {a:1,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
274. ~r2_hidden(X2,k10_xtuple_0(X3)) | ~v1_xboole_0(X4) | ~v5_relat_1(X3,X4) | ~v1_relat_1(X3) [resolution 263,106] {a:2,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
280. ~r2_hidden(X2,k10_xtuple_0(X3)) | m1_subset_1(X2,X4) | ~v5_relat_1(X3,X4) | ~v1_relat_1(X3) [resolution 265,106] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
305. ~v1_xboole_0(X0) | ~v5_relat_1(sK1,X0) | ~v1_relat_1(sK1) <- (2) [resolution 274,134] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
307. ~v5_relat_1(sK1,X0) | ~v1_xboole_0(X0) <- (2, 4) [subsumption resolution 305,144] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
308. ~v1_xboole_0(sK0) <- (2, 3, 4) [resolution 307,139] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
311. 25 <=> v1_xboole_0(sK0) [avatar definition]
313. ~v1_xboole_0(sK0) <- (~25) [avatar component clause 311] {a:4,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
314. ~25 | ~2 | ~3 | ~4 [avatar split clause 308,142,137,132,311]
315. m1_subset_1(sK2,X0) | ~v5_relat_1(sK1,X0) | ~v1_relat_1(sK1) <- (2) [resolution 280,134] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
317. ~v5_relat_1(sK1,X0) | m1_subset_1(sK2,X0) <- (2, 4) [subsumption resolution 315,144] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
318. m1_subset_1(sK2,sK0) <- (2, 3, 4) [resolution 317,139] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
321. 26 <=> m1_subset_1(sK2,sK0) [avatar definition]
323. m1_subset_1(sK2,sK0) <- (26) [avatar component clause 321] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
324. 26 | ~2 | ~3 | ~4 [avatar split clause 318,142,137,132,321]
333. v1_xboole_0(sK0) | r2_hidden(sK2,sK0) <- (26) [resolution 323,109] {a:5,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
335. r2_hidden(sK2,sK0) <- (~25, 26) [subsumption resolution 333,313] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
336. $false <- (~1, ~25, 26) [subsumption resolution 335,129] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
337. 1 | 25 | ~26 [avatar contradiction clause 336]
338. $false [avatar sat refutation 130,135,140,145,314,324,337] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:154,thDist:-154}
% SZS output end Proof for relat_1__t167_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 34
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 37
% Generated clauses: 143
% Active clauses: 69
% Passive clauses: 84
% Final active clauses: 68
% Final passive clauses: 12
% Main loop iterations started: 72
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 18
% 
% >>> Generating Inferences
% Binary resolution: 67
% Unit resulting resolution: 1
% Forward superposition: 1
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 29
% SAT solver unit clauses: 14
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% SAT solver time: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

