i: [1,117,0,16,3,1,0,0] t9_e_siec
a: [3,167,0,6,1,117]
s: 167
i: [1,118,0,2,1,0,0,1] fc1_xboole_0
a: [3,180,0,2,1,118]
s: 180
i: [1,151,0,2,1,0,0,2] rc2_xboole_0
a: [3,185,0,2,1,151]
s: 185
i: [1,152,0,2,1,0,0,2] rc1_xboole_0
a: [3,190,0,2,1,152]
s: 190
i: [1,153,0,2,1,0,0,2] existence_l1_e_siec
a: [3,195,0,2,1,153]
s: 195
i: [1,154,0,2,1,0,0,2] existence_l1_e_siec
a: [3,200,0,2,1,154]
s: 200
i: [1,156,0,2,1,0,0,2] rc1_relat_1
a: [3,205,0,2,1,156]
s: 205
i: [1,155,0,2,1,0,0,2] rc1_relat_1
a: [3,210,0,2,1,155]
s: 210
i: [1,157,0,2,1,0,0,255] rc2_relat_1
a: [3,215,0,2,1,157]
s: 215
i: [1,158,0,2,1,0,0,2] rc2_e_siec
a: [3,220,0,2,1,158]
s: 220
i: [1,159,0,2,1,0,0,2] rc3_e_siec
a: [3,225,0,2,1,159]
s: 225
i: [1,161,0,2,1,0,0,2] rc1_e_siec
a: [3,230,0,2,1,161]
s: 230
i: [1,160,0,2,1,0,0,2] rc1_e_siec
a: [3,235,0,2,1,160]
s: 235
i: [1,163,0,2,1,0,0,2] rc4_e_siec
a: [3,240,0,2,1,163]
s: 240
i: [1,162,0,2,1,0,0,2] rc4_e_siec
a: [3,245,0,2,1,162]
s: 245
i: [1,119,0,3,1,0,0,1] fc2_xboole_0
s: 119
i: [1,121,0,3,1,0,0,1] dt_k5_e_siec
s: 121
i: [1,122,0,3,1,0,0,1] dt_k5_e_siec
s: 122
i: [1,123,0,3,1,0,0,1] dt_k5_e_siec
s: 123
i: [1,124,0,3,1,0,0,1] dt_k5_e_siec
s: 124
i: [1,125,0,4,2,0,0,2] cc1_relat_1
s: 125
d: [2,247,1,2,1,1,40,125,180]
a: [3,252,1,2,1,247]
s: 252
i: [1,139,0,4,1,0,0,255] existence_m1_subset_1
s: 139
d: [2,248,1,2,1,1,40,125,190]
a: [3,257,1,2,1,248]
s: 257
i: [1,126,0,5,2,0,0,1] t6_boole
s: 126
i: [1,127,0,5,2,0,0,1] dt_u2_e_siec
s: 127
i: [1,128,0,5,2,0,0,1] dt_u2_e_siec
s: 128
i: [1,129,0,5,2,0,0,1] dt_u1_e_siec
s: 129
i: [1,130,0,5,2,0,0,1] dt_u1_e_siec
s: 130
i: [1,138,0,5,2,0,0,255] t7_boole
s: 138
i: [1,135,0,6,2,0,0,255] t1_subset
s: 135
i: [1,140,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 140
i: [1,137,0,7,3,0,0,1] t8_boole
s: 137
d: [2,259,1,7,3,0,44,126,126]
s: 259
i: [1,120,0,8,1,0,0,1] d7_e_siec
s: 120
i: [1,136,0,8,3,0,0,255] t2_subset
s: 136
d: [2,263,1,6,2,0,40,136,139]
s: 263
i: [1,141,0,9,3,0,0,2] dt_g1_e_siec
s: 141
d: [2,264,2,6,2,0,40,263,140]
s: 264
i: [1,142,0,9,3,0,0,2] dt_g1_e_siec
s: 142
i: [1,143,0,9,3,0,0,2] dt_g1_e_siec
s: 143
i: [1,144,0,9,3,0,0,2] dt_g1_e_siec
s: 144
i: [1,131,0,13,3,0,0,1] abstractness_v1_e_siec
s: 131
d: [2,262,1,10,2,0,44,120,126]
s: 262
i: [1,132,0,13,3,0,0,1] abstractness_v1_e_siec
s: 132
i: [1,145,0,16,4,0,0,2] free_g1_e_siec
s: 145
d: [2,278,2,17,5,0,44,145,262]
d: [2,284,2,15,4,0,18,278]
d: [2,287,2,13,3,0,21,284,125]
a: [3,290,2,2,1,287]
s: 290
r: 210
r: 185
r: 119
r: 125
r: 126
r: 138
r: 137
r: 262
d: [2,303,3,0,0,2,40,290,190]
e: 303
d: [2,302,3,0,0,2,40,290,180]
e: 302
d: [2,301,3,0,0,2,52,180,290]
e: 301
r: 290
s: 185
s: 210
s: 119
s: 125
s: 126
s: 138
s: 137
i: [1,146,0,16,4,0,0,2] free_g1_e_siec
s: 146
i: [1,147,0,16,4,0,0,2] free_g1_e_siec
s: 147
d: [2,309,1,7,3,0,44,126,126]
s: 309
i: [1,148,0,16,4,0,0,2] free_g1_e_siec
s: 148
s: 262
i: [1,149,0,16,4,0,0,2] free_g1_e_siec
s: 149
i: [1,150,0,16,4,0,0,2] free_g1_e_siec
s: 150
d: [2,310,1,10,2,0,44,120,126]
s: 310
d: [2,314,1,14,4,0,44,146,120]
d: [2,319,1,12,3,0,18,314]
d: [2,320,1,10,2,1,21,319,252]
s: 320
d: [2,323,1,14,4,0,44,147,120]
d: [2,328,1,12,3,0,18,323]
d: [2,329,1,10,2,1,21,328,252]
s: 329
d: [2,277,1,15,4,0,44,145,120]
d: [2,285,1,13,3,0,18,277]
d: [2,286,1,11,2,1,21,285,252]
s: 286
d: [2,279,1,20,6,0,44,145,131]
d: [2,295,1,17,5,0,21,279,127]
d: [2,296,1,14,4,0,21,295,129]
s: 296
d: [2,357,2,0,0,1,52,122,121,167,120,296]
e: 357
r: 167
a: [3,171,0,5,1,117]
s: 171
a: [3,166,0,6,1,117]
s: 166
d: [2,367,1,18,3,1,44,131,166]
d: [2,368,1,15,2,1,21,367,121]
d: [2,369,1,12,1,1,21,368,122]
a: [3,373,1,12,1,369]
s: 373
d: [2,375,2,0,0,3,52,171,373,320]
e: 375
r: 171
a: [3,170,0,5,1,117]
s: 170
r: 373
a: [3,175,0,5,1,117]
s: 175
d: [2,382,2,20,4,1,44,147,373]
a: [3,396,2,4,1,382]
s: 396
d: [2,415,1,10,1,2,23,373,170]
a: [3,421,1,10,1,415]
s: 421
d: [2,428,2,0,0,3,52,175,421,329]
e: 428
f: 428,415,375,369,357,247,118,117
% Refutation found. Thanks to Tanya!
% SZS status Theorem for e_siec__t9_e_siec
% SZS output start Proof for e_siec__t9_e_siec
1. ! [X0] : (k1_xboole_0 = u2_e_siec(k5_e_siec(X0)) & u1_e_siec(k5_e_siec(X0)) = k1_xboole_0 & u1_struct_0(k5_e_siec(X0)) = k1_tarski(X0)) [input t9_e_siec]
2. ~! [X0] : (k1_xboole_0 = u2_e_siec(k5_e_siec(X0)) & u1_e_siec(k5_e_siec(X0)) = k1_xboole_0 & u1_struct_0(k5_e_siec(X0)) = k1_tarski(X0)) [negated conjecture 1]
9. ! [X0,X1,X2] : ((v1_relat_1(X2) & v1_relat_1(X1)) => ! [X3,X4,X5] : (g1_e_siec(X0,X1,X2) = g1_e_siec(X3,X4,X5) => (X2 = X5 & X1 = X4 & X0 = X3))) [input free_g1_e_siec]
11. ! [X0] : (l1_e_siec(X0) => (v1_e_siec(X0) => g1_e_siec(u1_struct_0(X0),u1_e_siec(X0),u2_e_siec(X0)) = X0)) [input abstractness_v1_e_siec]
21. ! [X0] : (v1_xboole_0(X0) => v1_relat_1(X0)) [input cc1_relat_1]
36. ! [X0] : (l1_e_siec(k5_e_siec(X0)) & v3_e_siec(k5_e_siec(X0)) & v2_e_siec(k5_e_siec(X0)) & v1_e_siec(k5_e_siec(X0))) [input dt_k5_e_siec]
38. ! [X0] : (l1_e_siec(X0) => v1_relat_1(u1_e_siec(X0))) [input dt_u1_e_siec]
41. ! [X0] : (l1_e_siec(X0) => v1_relat_1(u2_e_siec(X0))) [input dt_u2_e_siec]
42. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
45. ! [X0] : k5_e_siec(X0) = g1_e_siec(k1_tarski(X0),k1_xboole_0,k1_xboole_0) [input d7_e_siec]
53. ! [X0] : (l1_e_siec(k5_e_siec(X0)) & v2_e_siec(k5_e_siec(X0)) & v1_e_siec(k5_e_siec(X0))) [pure predicate removal 36]
56. ! [X0] : (l1_e_siec(k5_e_siec(X0)) & v1_e_siec(k5_e_siec(X0))) [pure predicate removal 53]
64. ? [X0] : (k1_xboole_0 != u2_e_siec(k5_e_siec(X0)) | u1_e_siec(k5_e_siec(X0)) != k1_xboole_0 | u1_struct_0(k5_e_siec(X0)) != k1_tarski(X0)) [ennf transformation 2]
65. ! [X0] : (v1_relat_1(X0) | ~v1_xboole_0(X0)) [ennf transformation 21]
67. ! [X0] : (v1_relat_1(u2_e_siec(X0)) | ~l1_e_siec(X0)) [ennf transformation 41]
69. ! [X0] : (v1_relat_1(u1_e_siec(X0)) | ~l1_e_siec(X0)) [ennf transformation 38]
71. ! [X0] : ((g1_e_siec(u1_struct_0(X0),u1_e_siec(X0),u2_e_siec(X0)) = X0 | ~v1_e_siec(X0)) | ~l1_e_siec(X0)) [ennf transformation 11]
72. ! [X0] : (g1_e_siec(u1_struct_0(X0),u1_e_siec(X0),u2_e_siec(X0)) = X0 | ~v1_e_siec(X0) | ~l1_e_siec(X0)) [flattening 71]
89. ! [X0,X1,X2] : (! [X3,X4,X5] : ((X2 = X5 & X1 = X4 & X0 = X3) | g1_e_siec(X0,X1,X2) != g1_e_siec(X3,X4,X5)) | (~v1_relat_1(X2) | ~v1_relat_1(X1))) [ennf transformation 9]
90. ! [X0,X1,X2] : (! [X3,X4,X5] : ((X2 = X5 & X1 = X4 & X0 = X3) | g1_e_siec(X0,X1,X2) != g1_e_siec(X3,X4,X5)) | ~v1_relat_1(X2) | ~v1_relat_1(X1)) [flattening 89]
93. ? [X0] : (k1_xboole_0 != u2_e_siec(k5_e_siec(X0)) | u1_e_siec(k5_e_siec(X0)) != k1_xboole_0 | u1_struct_0(k5_e_siec(X0)) != k1_tarski(X0)) => (k1_xboole_0 != u2_e_siec(k5_e_siec(sK0)) | k1_xboole_0 != u1_e_siec(k5_e_siec(sK0)) | u1_struct_0(k5_e_siec(sK0)) != k1_tarski(sK0)) [choice axiom]
94. k1_xboole_0 != u2_e_siec(k5_e_siec(sK0)) | k1_xboole_0 != u1_e_siec(k5_e_siec(sK0)) | u1_struct_0(k5_e_siec(sK0)) != k1_tarski(sK0) [skolemisation 64,93]
117. k1_xboole_0 != u2_e_siec(k5_e_siec(sK0)) | k1_xboole_0 != u1_e_siec(k5_e_siec(sK0)) | u1_struct_0(k5_e_siec(sK0)) != k1_tarski(sK0) [cnf transformation 94] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
118. v1_xboole_0(k1_xboole_0) [cnf transformation 42] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
120. k5_e_siec(X0) = g1_e_siec(k1_tarski(X0),k1_xboole_0,k1_xboole_0) [cnf transformation 45] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
121. v1_e_siec(k5_e_siec(X0)) [cnf transformation 56] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
122. l1_e_siec(k5_e_siec(X0)) [cnf transformation 56] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
125. ~v1_xboole_0(X0) | v1_relat_1(X0) [cnf transformation 65] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
127. v1_relat_1(u2_e_siec(X0)) | ~l1_e_siec(X0) [cnf transformation 67] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
129. v1_relat_1(u1_e_siec(X0)) | ~l1_e_siec(X0) [cnf transformation 69] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
131. g1_e_siec(u1_struct_0(X0),u1_e_siec(X0),u2_e_siec(X0)) = X0 | ~v1_e_siec(X0) | ~l1_e_siec(X0) [cnf transformation 72] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
145. g1_e_siec(X0,X1,X2) != g1_e_siec(X3,X4,X5) | X0 = X3 | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 90] {a:0,w:16,wCS:6400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
146. g1_e_siec(X0,X1,X2) != g1_e_siec(X3,X4,X5) | X1 = X4 | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 90] {a:0,w:16,wCS:6400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
147. g1_e_siec(X0,X1,X2) != g1_e_siec(X3,X4,X5) | X2 = X5 | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 90] {a:0,w:16,wCS:6400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
165. 1 <=> u1_struct_0(k5_e_siec(sK0)) = k1_tarski(sK0) [avatar definition]
166. u1_struct_0(k5_e_siec(sK0)) = k1_tarski(sK0) <- (1) [avatar component clause 165] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
167. u1_struct_0(k5_e_siec(sK0)) != k1_tarski(sK0) <- (~1) [avatar component clause 165] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
169. 2 <=> k1_xboole_0 = u1_e_siec(k5_e_siec(sK0)) [avatar definition]
170. k1_xboole_0 = u1_e_siec(k5_e_siec(sK0)) <- (2) [avatar component clause 169] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
171. k1_xboole_0 != u1_e_siec(k5_e_siec(sK0)) <- (~2) [avatar component clause 169] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
173. 3 <=> k1_xboole_0 = u2_e_siec(k5_e_siec(sK0)) [avatar definition]
175. k1_xboole_0 != u2_e_siec(k5_e_siec(sK0)) <- (~3) [avatar component clause 173] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
176. ~1 | ~2 | ~3 [avatar split clause 117,173,169,165]
178. 4 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
180. v1_xboole_0(k1_xboole_0) <- (4) [avatar component clause 178] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
181. 4 [avatar split clause 118,178]
247. v1_relat_1(k1_xboole_0) <- (4) [resolution 125,180] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
250. 18 <=> v1_relat_1(k1_xboole_0) [avatar definition]
252. v1_relat_1(k1_xboole_0) <- (18) [avatar component clause 250] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
253. 18 | ~4 [avatar split clause 247,178,250]
279. g1_e_siec(X10,X11,X12) != X9 | u1_struct_0(X9) = X10 | ~v1_relat_1(u2_e_siec(X9)) | ~v1_relat_1(u1_e_siec(X9)) | ~v1_e_siec(X9) | ~l1_e_siec(X9) [superposition 145,131] {a:1,w:20,wCS:8000,sine:1,thAx:0,allAx:2,thDist:-2}
295. g1_e_siec(X10,X11,X12) != X9 | u1_struct_0(X9) = X10 | ~v1_relat_1(u1_e_siec(X9)) | ~v1_e_siec(X9) | ~l1_e_siec(X9) [subsumption resolution 279,127] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
296. g1_e_siec(X10,X11,X12) != X9 | u1_struct_0(X9) = X10 | ~v1_e_siec(X9) | ~l1_e_siec(X9) [subsumption resolution 295,129] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
314. k5_e_siec(X0) != g1_e_siec(X1,X2,X3) | k1_xboole_0 = X2 | ~v1_relat_1(k1_xboole_0) | ~v1_relat_1(k1_xboole_0) [superposition 146,120] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
319. k5_e_siec(X0) != g1_e_siec(X1,X2,X3) | k1_xboole_0 = X2 | ~v1_relat_1(k1_xboole_0) [duplicate literal removal 314] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
320. k5_e_siec(X0) != g1_e_siec(X1,X2,X3) | k1_xboole_0 = X2 <- (18) [subsumption resolution 319,252] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
323. k5_e_siec(X0) != g1_e_siec(X1,X2,X3) | k1_xboole_0 = X3 | ~v1_relat_1(k1_xboole_0) | ~v1_relat_1(k1_xboole_0) [superposition 147,120] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
328. k5_e_siec(X0) != g1_e_siec(X1,X2,X3) | k1_xboole_0 = X3 | ~v1_relat_1(k1_xboole_0) [duplicate literal removal 323] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
329. k5_e_siec(X0) != g1_e_siec(X1,X2,X3) | k1_xboole_0 = X3 <- (18) [subsumption resolution 328,252] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
357. $false <- (~1) [unit resulting resolution 122,121,167,120,296] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
362. 1 [avatar contradiction clause 357]
367. k5_e_siec(sK0) = g1_e_siec(k1_tarski(sK0),u1_e_siec(k5_e_siec(sK0)),u2_e_siec(k5_e_siec(sK0))) | ~v1_e_siec(k5_e_siec(sK0)) | ~l1_e_siec(k5_e_siec(sK0)) <- (1) [superposition 131,166] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
368. k5_e_siec(sK0) = g1_e_siec(k1_tarski(sK0),u1_e_siec(k5_e_siec(sK0)),u2_e_siec(k5_e_siec(sK0))) | ~l1_e_siec(k5_e_siec(sK0)) <- (1) [subsumption resolution 367,121] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
369. k5_e_siec(sK0) = g1_e_siec(k1_tarski(sK0),u1_e_siec(k5_e_siec(sK0)),u2_e_siec(k5_e_siec(sK0))) <- (1) [subsumption resolution 368,122] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
371. 24 <=> k5_e_siec(sK0) = g1_e_siec(k1_tarski(sK0),u1_e_siec(k5_e_siec(sK0)),u2_e_siec(k5_e_siec(sK0))) [avatar definition]
373. k5_e_siec(sK0) = g1_e_siec(k1_tarski(sK0),u1_e_siec(k5_e_siec(sK0)),u2_e_siec(k5_e_siec(sK0))) <- (24) [avatar component clause 371] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
374. 24 | ~1 [avatar split clause 369,165,371]
375. $false <- (~2, 18, 24) [unit resulting resolution 171,373,320] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
387. 2 | ~18 | ~24 [avatar contradiction clause 375]
415. k5_e_siec(sK0) = g1_e_siec(k1_tarski(sK0),k1_xboole_0,u2_e_siec(k5_e_siec(sK0))) <- (2, 24) [backward demodulation 373,170] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
419. 31 <=> k5_e_siec(sK0) = g1_e_siec(k1_tarski(sK0),k1_xboole_0,u2_e_siec(k5_e_siec(sK0))) [avatar definition]
421. k5_e_siec(sK0) = g1_e_siec(k1_tarski(sK0),k1_xboole_0,u2_e_siec(k5_e_siec(sK0))) <- (31) [avatar component clause 419] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
422. 31 | ~2 | ~24 [avatar split clause 415,371,169,419]
428. $false <- (~3, 18, 31) [unit resulting resolution 175,421,329] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
441. 3 | ~18 | ~31 [avatar contradiction clause 428]
445. $false [avatar sat refutation 176,181,253,362,374,387,422,441] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:179,thDist:-179}
% SZS output end Proof for e_siec__t9_e_siec
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 44
% >>> Preprocessing
% Introduced skolems: 12
% Pure predicates: 5
% 
% >>> Saturation
% Initial clauses: 47
% Generated clauses: 222
% Active clauses: 58
% Passive clauses: 99
% Final active clauses: 46
% Final passive clauses: 26
% Main loop iterations started: 73
% 
% >>> Simplifying Inferences
% Duplicate literals: 13
% Fw subsumption resolutions: 21
% Fw demodulations: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Equational tautologies: 3
% Forward subsumptions: 34
% Backward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 4
% Forward superposition: 25
% Backward superposition: 41
% Self superposition: 2
% Equality resolution: 4
% 
% >>> AVATAR
% Split clauses: 11
% Split components: 26
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 40
% SAT solver unit clauses: 15
% SAT solver binary clauses: 12
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% property evaluation: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

