i: [1,90,0,2,1,1,0,0] t75_semi_af1
a: [3,144,0,2,1,90]
s: 144
i: [1,89,0,2,1,1,0,0] t75_semi_af1
a: [3,149,0,2,1,89]
s: 149
i: [1,88,0,2,1,1,0,0] t75_semi_af1
a: [3,154,0,2,1,88]
s: 154
i: [1,92,0,4,1,1,0,0] t75_semi_af1
a: [3,134,0,4,1,92]
s: 134
i: [1,91,0,4,1,1,0,0] t75_semi_af1
a: [3,139,0,4,1,91]
s: 139
i: [1,93,0,7,1,1,0,0] t75_semi_af1
a: [3,129,0,7,1,93]
s: 129
i: [1,120,0,2,1,0,0,2] existence_l1_struct_0
a: [3,159,0,2,1,120]
s: 159
i: [1,121,0,2,1,0,0,1] existence_l1_analoaf
a: [3,164,0,2,1,121]
s: 164
i: [1,124,0,2,1,0,0,1] rc1_semi_af1
a: [3,169,0,2,1,124]
s: 169
i: [1,123,0,2,1,0,0,1] rc1_semi_af1
a: [3,174,0,2,1,123]
s: 174
i: [1,122,0,2,1,0,0,1] rc1_semi_af1
a: [3,179,0,2,1,122]
s: 179
i: [1,108,0,4,2,0,0,1] dt_l1_analoaf
s: 108
i: [1,118,0,4,1,0,0,1] existence_m1_subset_1
s: 118
d: [2,181,1,2,1,1,40,108,144]
a: [3,187,1,2,1,181]
s: 187
i: [1,94,0,6,3,0,0,255] cc7_struct_0
s: 94
i: [1,95,0,6,3,0,0,255] cc5_struct_0
s: 95
d: [2,182,1,2,1,1,40,108,164]
a: [3,192,1,2,1,182]
s: 192
i: [1,97,0,6,3,0,0,255] cc17_struct_0
s: 97
d: [2,183,1,2,1,1,40,108,179]
a: [3,197,1,2,1,183]
s: 197
i: [1,98,0,6,3,0,0,255] cc2_struct_0
s: 98
i: [1,99,0,6,3,0,0,255] cc6_struct_0
s: 99
d: [2,199,1,4,2,1,40,95,159]
a: [3,204,1,2,1,199]
s: 204
i: [1,100,0,6,3,0,0,255] cc1_struct_0
s: 100
d: [2,210,2,4,2,1,40,192,95]
a: [3,214,2,2,1,210]
s: 214
i: [1,103,0,6,3,0,0,255] cc4_struct_0
s: 103
i: [1,101,0,7,3,0,0,255] cc8_struct_0
s: 101
i: [1,104,0,7,3,0,0,255] cc9_struct_0
s: 104
i: [1,105,0,7,3,0,0,255] cc11_struct_0
s: 105
i: [1,106,0,7,3,0,0,255] cc11_struct_0
s: 106
i: [1,109,0,7,3,0,0,255] fc9_struct_0
s: 109
i: [1,110,0,7,3,0,0,255] fc6_struct_0
s: 110
i: [1,114,0,7,3,0,0,255] fc2_struct_0
s: 114
i: [1,115,0,7,3,0,0,255] fc8_struct_0
s: 115
i: [1,116,0,7,3,0,0,255] fc7_struct_0
s: 116
i: [1,117,0,7,3,0,0,255] fc1_struct_0
s: 117
i: [1,107,0,9,4,0,0,255] cc10_struct_0
s: 107
i: [1,111,0,20,6,0,0,2] t64_semi_af1
s: 111
i: [1,119,0,26,7,0,0,1] dt_k1_semi_af1
s: 119
i: [1,113,0,35,9,0,0,1] d5_semi_af1
s: 113
i: [1,125,0,36,8,0,0,1] d5_semi_af1
s: 125
d: [2,258,0,28,7,0,21,125,119]
s: 258
d: [2,248,1,31,9,1,44,129,113]
d: [2,251,1,27,8,1,18,248]
d: [2,253,1,23,7,2,21,251,134]
d: [2,254,1,19,6,3,21,253,139]
d: [2,255,1,17,5,4,21,254,144]
d: [2,256,1,15,4,5,21,255,149]
d: [2,257,1,13,3,6,21,256,154]
s: 257
d: [2,247,1,59,17,0,44,113,113]
d: [2,252,1,41,11,0,18,247]
s: 252
d: [2,261,2,10,2,6,47,257]
d: [2,262,2,6,1,6,21,261,139]
a: [3,266,2,6,1,262]
s: 266
d: [2,272,3,0,0,6,52,154,149,144,134,139,266,111]
e: 272
d: [2,273,3,14,5,1,40,266,111]
d: [2,275,3,10,4,2,21,273,139]
d: [2,276,3,6,3,3,21,275,134]
d: [2,277,3,4,2,4,21,276,144]
d: [2,278,3,2,1,5,21,277,149]
d: [2,279,3,0,0,6,21,278,154]
e: 279
f: 279,262,88,89,90,91,92,93
% Refutation found. Thanks to Tanya!
% SZS status Theorem for semi_af1__t75_semi_af1
% SZS output start Proof for semi_af1__t75_semi_af1
1. ! [X0] : ((l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => k1_semi_af1(X0,X1,X2,X2) = X1))) [input t75_semi_af1]
2. ~! [X0] : ((l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => k1_semi_af1(X0,X1,X2,X2) = X1))) [negated conjecture 1]
14. ! [X0] : ((l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (k1_semi_af1(X0,X1,X2,X3) = X4 <=> r3_semi_af1(X0,X3,X1,X2,X4))))))) [input d5_semi_af1]
31. ! [X0] : ((l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => r3_semi_af1(X0,X1,X2,X1,X2)))) [input t64_semi_af1]
32. ? [X0] : (? [X1] : (? [X2] : (k1_semi_af1(X0,X1,X2,X2) != X1 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
33. ? [X0] : (? [X1] : (? [X2] : (k1_semi_af1(X0,X1,X2,X2) != X1 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) [flattening 32]
61. ! [X0] : (! [X1] : (! [X2] : (r3_semi_af1(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0))) [ennf transformation 31]
62. ! [X0] : (! [X1] : (! [X2] : (r3_semi_af1(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0)) [flattening 61]
63. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((k1_semi_af1(X0,X1,X2,X3) = X4 <=> r3_semi_af1(X0,X3,X1,X2,X4)) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0))) [ennf transformation 14]
64. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((k1_semi_af1(X0,X1,X2,X3) = X4 <=> r3_semi_af1(X0,X3,X1,X2,X4)) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0)) [flattening 63]
75. ? [X0] : (? [X1] : (? [X2] : (k1_semi_af1(X0,X1,X2,X2) != X1 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_semi_af1(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (k1_semi_af1(sK0,X1,X2,X2) != X1 & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_semi_af1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
76. ? [X1] : (? [X2] : (k1_semi_af1(sK0,X1,X2,X2) != X1 & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (sK1 != k1_semi_af1(sK0,sK1,X2,X2) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
77. ? [X2] : (sK1 != k1_semi_af1(sK0,sK1,X2,X2) & m1_subset_1(X2,u1_struct_0(sK0))) => (sK1 != k1_semi_af1(sK0,sK1,sK2,sK2) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
78. ((sK1 != k1_semi_af1(sK0,sK1,sK2,sK2) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_semi_af1(sK0) & ~v2_struct_0(sK0) [skolemisation 33,77,76,75]
79. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((k1_semi_af1(X0,X1,X2,X3) = X4 | ~r3_semi_af1(X0,X3,X1,X2,X4)) & (r3_semi_af1(X0,X3,X1,X2,X4) | k1_semi_af1(X0,X1,X2,X3) != X4)) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0)) [nnf transformation 64]
88. ~v2_struct_0(sK0) [cnf transformation 78] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
89. v1_semi_af1(sK0) [cnf transformation 78] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
90. l1_analoaf(sK0) [cnf transformation 78] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
91. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 78] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
92. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 78] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
93. sK1 != k1_semi_af1(sK0,sK1,sK2,sK2) [cnf transformation 78] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
111. r3_semi_af1(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0) [cnf transformation 62] {a:0,w:20,wCS:8000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
113. k1_semi_af1(X0,X1,X2,X3) = X4 | ~r3_semi_af1(X0,X3,X1,X2,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_semi_af1(X0) | v2_struct_0(X0) [cnf transformation 79] {a:0,w:35,wCS:14000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
127. 1 <=> sK1 = k1_semi_af1(sK0,sK1,sK2,sK2) [avatar definition]
129. sK1 != k1_semi_af1(sK0,sK1,sK2,sK2) <- (~1) [avatar component clause 127] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
130. ~1 [avatar split clause 93,127]
132. 2 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
134. m1_subset_1(sK2,u1_struct_0(sK0)) <- (2) [avatar component clause 132] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
135. 2 [avatar split clause 92,132]
137. 3 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
139. m1_subset_1(sK1,u1_struct_0(sK0)) <- (3) [avatar component clause 137] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
140. 3 [avatar split clause 91,137]
142. 4 <=> l1_analoaf(sK0) [avatar definition]
144. l1_analoaf(sK0) <- (4) [avatar component clause 142] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
145. 4 [avatar split clause 90,142]
147. 5 <=> v1_semi_af1(sK0) [avatar definition]
149. v1_semi_af1(sK0) <- (5) [avatar component clause 147] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
150. 5 [avatar split clause 89,147]
152. 6 <=> v2_struct_0(sK0) [avatar definition]
154. ~v2_struct_0(sK0) <- (~6) [avatar component clause 152] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
155. ~6 [avatar split clause 88,152]
248. sK1 != X0 | ~r3_semi_af1(sK0,sK2,sK1,sK2,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~1) [superposition 129,113] {a:1,w:31,wCS:3100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
251. sK1 != X0 | ~r3_semi_af1(sK0,sK2,sK1,sK2,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~1) [duplicate literal removal 248] {a:1,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
253. sK1 != X0 | ~r3_semi_af1(sK0,sK2,sK1,sK2,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~1, 2) [subsumption resolution 251,134] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
254. sK1 != X0 | ~r3_semi_af1(sK0,sK2,sK1,sK2,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~1, 2, 3) [subsumption resolution 253,139] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
255. sK1 != X0 | ~r3_semi_af1(sK0,sK2,sK1,sK2,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~1, 2, 3, 4) [subsumption resolution 254,144] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
256. sK1 != X0 | ~r3_semi_af1(sK0,sK2,sK1,sK2,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~1, 2, 3, 4, 5) [subsumption resolution 255,149] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
257. sK1 != X0 | ~r3_semi_af1(sK0,sK2,sK1,sK2,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 2, 3, 4, 5, ~6) [subsumption resolution 256,154] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
261. ~r3_semi_af1(sK0,sK2,sK1,sK2,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) <- (~1, 2, 3, 4, 5, ~6) [equality resolution 257] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
262. ~r3_semi_af1(sK0,sK2,sK1,sK2,sK1) <- (~1, 2, 3, 4, 5, ~6) [subsumption resolution 261,139] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
264. 21 <=> r3_semi_af1(sK0,sK2,sK1,sK2,sK1) [avatar definition]
266. ~r3_semi_af1(sK0,sK2,sK1,sK2,sK1) <- (~21) [avatar component clause 264] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
267. ~21 | 1 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 262,152,147,142,137,132,127,264]
273. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (~21) [resolution 266,111] {a:3,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
275. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (3, ~21) [subsumption resolution 273,139] {a:3,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
276. ~l1_analoaf(sK0) | ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (2, 3, ~21) [subsumption resolution 275,134] {a:3,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
277. ~v1_semi_af1(sK0) | v2_struct_0(sK0) <- (2, 3, 4, ~21) [subsumption resolution 276,144] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
278. v2_struct_0(sK0) <- (2, 3, 4, 5, ~21) [subsumption resolution 277,149] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
279. $false <- (2, 3, 4, 5, ~6, ~21) [subsumption resolution 278,154] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
280. ~2 | ~3 | ~4 | ~5 | 6 | 21 [avatar contradiction clause 279]
281. $false [avatar sat refutation 130,135,140,145,150,155,267,280] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:132,thDist:-132}
% SZS output end Proof for semi_af1__t75_semi_af1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 30
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 37
% Generated clauses: 104
% Active clauses: 39
% Passive clauses: 45
% Final active clauses: 39
% Final passive clauses: 1
% Main loop iterations started: 44
% 
% >>> Simplifying Inferences
% Duplicate literals: 38
% Fw subsumption resolutions: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 1
% Forward superposition: 1
% Backward superposition: 2
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 11
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

