i: [1,65634,0,32,16,1,0,0] t1_gate_2
a: [3,65661,0,2,1,65634]
s: 65661
a: [3,65666,0,2,1,65634]
s: 65666
a: [3,65669,0,2,1,65634]
s: 65669
a: [3,65673,0,2,1,65634]
s: 65673
a: [3,65678,0,2,1,65634]
s: 65678
a: [3,65681,0,2,1,65634]
s: 65681
a: [3,65686,0,2,1,65634]
s: 65686
a: [3,65690,0,2,1,65634]
s: 65690
a: [3,65693,0,2,1,65634]
s: 65693
a: [3,65698,0,2,1,65634]
s: 65698
a: [3,65702,0,2,1,65634]
s: 65702
a: [3,65705,0,2,1,65634]
s: 65705
a: [3,65710,0,2,1,65634]
s: 65710
a: [3,65713,0,2,1,65634]
s: 65713
a: [3,65718,0,2,1,65634]
s: 65718
a: [3,65722,0,2,1,65634]
s: 65722
i: [1,98,0,13,2,1,0,0] t1_gate_2
a: [3,65986,0,2,1,98]
s: 65986
i: [1,96,0,6,2,1,0,0] t1_gate_2
a: [3,65996,0,2,1,96]
s: 65996
i: [1,94,0,5,2,1,0,0] t1_gate_2
a: [3,66005,0,2,1,94]
s: 66005
a: [3,66002,0,3,1,94]
s: 66002
a: [3,65991,0,4,1,96]
s: 65991
i: [1,92,0,7,2,1,0,0] t1_gate_2
a: [3,66011,0,5,1,92]
s: 66011
i: [1,76,0,7,2,1,0,0] t1_gate_2
a: [3,66060,0,5,1,76]
s: 66060
i: [1,90,0,8,2,1,0,0] t1_gate_2
a: [3,66017,0,6,1,90]
s: 66017
i: [1,88,0,8,2,1,0,0] t1_gate_2
a: [3,66024,0,6,1,88]
s: 66024
i: [1,84,0,8,2,1,0,0] t1_gate_2
a: [3,66035,0,6,1,84]
s: 66035
i: [1,74,0,8,2,1,0,0] t1_gate_2
a: [3,66066,0,6,1,74]
s: 66066
i: [1,72,0,8,2,1,0,0] t1_gate_2
a: [3,66072,0,6,1,72]
s: 66072
i: [1,68,0,8,2,1,0,0] t1_gate_2
a: [3,66084,0,6,1,68]
s: 66084
i: [1,86,0,9,2,1,0,0] t1_gate_2
a: [3,66030,0,7,1,86]
s: 66030
i: [1,82,0,9,2,1,0,0] t1_gate_2
a: [3,66041,0,7,1,82]
s: 66041
i: [1,80,0,9,2,1,0,0] t1_gate_2
a: [3,66048,0,7,1,80]
s: 66048
i: [1,70,0,9,2,1,0,0] t1_gate_2
a: [3,66077,0,7,1,70]
s: 66077
i: [1,66,0,9,2,1,0,0] t1_gate_2
a: [3,66090,0,7,1,66]
s: 66090
i: [1,64,0,9,2,1,0,0] t1_gate_2
a: [3,66095,0,7,1,64]
s: 66095
i: [1,78,0,10,2,1,0,0] t1_gate_2
a: [3,66054,0,8,1,78]
s: 66054
i: [1,62,0,10,2,1,0,0] t1_gate_2
a: [3,66101,0,8,1,62]
s: 66101
a: [3,65981,0,11,1,98]
s: 65981
i: [1,65635,0,5,2,0,0,1] fc2_gate_1
s: 65635
i: [1,65636,0,5,2,0,0,1] fc1_gate_1
s: 65636
i: [1,65637,0,5,2,0,0,255] t6_boole
s: 65637
d: [2,66105,1,2,1,1,40,65636,66002]
a: [3,66110,1,2,1,66105]
s: 66110
i: [1,65648,0,6,2,0,0,1] fc5_gate_1
s: 65648
i: [1,65649,0,6,2,0,0,1] fc4_gate_1
s: 65649
i: [1,65644,0,7,3,0,0,2] t8_boole
s: 65644
i: [1,65645,0,7,1,0,0,1] commutativity_k3_gate_1
s: 65645
d: [2,66113,1,6,2,0,44,65648,65645]
s: 66113
i: [1,65646,0,7,1,0,0,1] commutativity_k2_gate_1
s: 65646
d: [2,66114,1,6,2,0,44,65648,65645]
s: 66114
i: [1,65647,0,7,1,0,0,2] commutativity_k4_gate_1
s: 65647
i: [1,65655,0,7,2,0,0,1] fc18_gate_1
s: 65655
d: [2,66123,1,2,1,1,40,65655,66077]
a: [3,66135,1,2,1,66123]
s: 66135
i: [1,65656,0,7,2,0,0,1] fc19_gate_1
s: 65656
d: [2,66121,1,3,1,1,40,65655,66035]
a: [3,66129,1,3,1,66121]
s: 66129
d: [2,66157,2,0,0,2,52,65986,66129,65635]
e: 66157
d: [2,66158,2,2,1,1,40,66129,65635]
d: [2,66160,2,0,0,2,21,66158,65986]
e: 66160
r: 65981
r: 65986
r: 65673
r: 65681
r: 66035
r: 66041
r: 66129
a: [3,65674,0,2,1,65634]
s: 65674
a: [3,65682,0,2,1,65634]
s: 65682
s: 65986
a: [3,66036,0,6,1,84]
s: 66036
a: [3,66128,1,3,1,66121]
s: 66128
a: [3,66042,0,7,1,82]
s: 66042
s: 65981
d: [2,66124,1,3,1,1,40,65655,66101]
a: [3,66140,1,3,1,66124]
s: 66140
d: [2,66166,2,0,0,2,52,66135,66140,65635]
e: 66166
d: [2,66167,2,2,1,1,40,66140,65635]
d: [2,66169,2,0,0,2,21,66167,66135]
e: 66169
r: 65693
r: 66077
r: 66135
a: [3,65694,0,2,1,65634]
s: 65694
a: [3,66134,1,2,1,66123]
s: 66134
a: [3,66078,0,7,1,70]
s: 66078
i: [1,65657,0,7,2,0,0,1] fc20_gate_1
s: 65657
d: [2,66172,1,0,0,2,52,66005,66011,65657]
e: 66172
d: [2,66173,1,0,0,2,52,66002,66101,65657]
e: 66173
d: [2,66174,1,2,1,1,40,65657,66011]
d: [2,66180,1,0,0,2,21,66174,66005]
e: 66180
d: [2,66176,1,3,1,1,40,65657,66101]
d: [2,66182,1,0,0,2,21,66176,66002]
e: 66182
r: 66002
r: 66005
a: [3,66006,0,2,1,94]
s: 66006
a: [3,66001,0,3,1,94]
s: 66001
d: [2,66191,2,0,0,2,52,66110,66001,65635]
e: 66191
d: [2,66192,1,2,1,1,40,66001,65635]
d: [2,66194,1,0,0,2,21,66192,66110]
e: 66194
r: 66110
r: 65669
r: 66095
a: [3,65670,0,2,1,65634]
s: 65670
a: [3,66096,0,7,1,64]
s: 66096
i: [1,65639,0,8,3,0,0,1] t6_gate_1
s: 65639
a: [3,66109,1,2,1,66105]
s: 66109
i: [1,65640,0,8,3,0,0,1] t6_gate_1
s: 65640
d: [2,66198,1,4,2,1,40,65639,65991]
a: [3,66203,1,2,1,66198]
s: 66203
i: [1,65642,0,8,3,0,0,1] t6_gate_1
s: 65642
i: [1,65643,0,8,3,0,0,1] t6_gate_1
s: 65643
d: [2,66147,1,3,1,1,40,65656,66077]
a: [3,66153,1,3,1,66147]
s: 66153
d: [2,66223,2,0,0,2,52,66203,66153,65635]
e: 66223
d: [2,66224,2,2,1,1,40,66153,65635]
d: [2,66226,2,0,0,2,21,66224,66203]
e: 66226
r: 65991
r: 65996
r: 66203
r: 65661
r: 66101
r: 66153
a: [3,65662,0,2,1,65634]
s: 65662
s: 65996
s: 65991
a: [3,66102,0,8,1,62]
s: 66102
s: 66203
i: [1,65650,0,8,3,0,0,1] fc3_gate_1
s: 65650
a: [3,66152,1,3,1,66147]
s: 66152
i: [1,65651,0,8,3,0,0,1] fc9_gate_1
s: 65651
i: [1,65652,0,8,3,0,0,1] fc7_gate_1
s: 65652
d: [2,66175,1,3,1,1,40,65657,66017]
a: [3,66189,1,3,1,66175]
s: 66189
d: [2,66239,2,0,0,2,52,66006,66189,65635]
e: 66239
d: [2,66240,2,2,1,1,40,66189,65635]
d: [2,66242,2,0,0,2,21,66240,66006]
e: 66242
r: 65713
r: 66001
r: 66006
r: 66011
r: 66109
r: 66203
r: 66152
a: [3,65714,0,2,1,65634]
s: 65714
s: 66005
s: 66002
a: [3,66012,0,5,1,92]
s: 66012
s: 66110
i: [1,65653,0,8,3,0,0,1] fc6_gate_1
s: 65653
a: [3,66204,1,2,1,66198]
s: 66204
i: [1,65654,0,8,3,0,0,1] fc8_gate_1
s: 65654
i: [1,65658,0,11,4,0,0,1] fc17_gate_1
s: 65658
r: 65698
r: 65705
r: 65981
r: 65986
r: 66017
r: 66024
r: 66140
r: 66128
r: 65674
r: 65722
r: 66042
r: 66054
s: 65673
a: [3,65697,0,2,1,65634]
s: 65697
a: [3,65706,0,2,1,65634]
s: 65706
a: [3,65721,0,2,1,65634]
s: 65721
a: [3,65985,0,2,1,98]
s: 65985
s: 66129
a: [3,66018,0,6,1,90]
s: 66018
a: [3,66139,1,3,1,66124]
s: 66139
d: [2,66302,2,0,0,2,52,66134,66139,65636]
e: 66302
d: [2,66303,2,2,1,1,40,66139,65636]
d: [2,66305,2,0,0,2,21,66303,66134]
e: 66305
r: 66134
r: 65702
r: 66072
a: [3,65701,0,2,1,65634]
s: 65701
a: [3,66023,0,6,1,88]
s: 66023
d: [2,66310,1,0,0,2,52,66005,66023,65657]
e: 66310
d: [2,66311,1,0,0,2,52,65985,66023,65655]
e: 66311
d: [2,66312,1,2,1,1,40,66023,65657]
d: [2,66317,1,0,0,2,21,66312,66005]
e: 66317
d: [2,66314,1,2,1,1,40,66023,65655]
d: [2,66319,1,0,0,2,21,66314,65985]
e: 66319
r: 65697
r: 66023
s: 65698
s: 66024
s: 66135
a: [3,66071,0,6,1,72]
s: 66071
s: 66041
s: 66153
a: [3,66053,0,8,1,78]
s: 66053
d: [2,66280,1,8,3,1,40,65658,66030]
d: [2,66285,1,5,2,2,21,66280,66189]
d: [2,66286,1,3,1,3,21,66285,65986]
a: [3,66289,1,3,1,66286]
s: 66289
d: [2,66334,2,0,0,2,52,65996,66289,65635]
e: 66334
d: [2,66335,2,2,1,1,40,66289,65635]
d: [2,66337,2,0,0,2,21,66335,65996]
e: 66337
r: 65991
r: 65996
r: 65673
r: 66041
s: 65674
a: [3,65995,0,2,1,96]
s: 65995
a: [3,65992,0,4,1,96]
s: 65992
d: [2,66341,2,0,0,3,52,66110,66204,65992,65642]
e: 66341
d: [2,66343,1,4,2,1,40,65992,65642]
d: [2,66345,1,2,1,2,21,66343,66204]
d: [2,66346,1,0,0,3,21,66345,66110]
e: 66346
r: 65701
r: 66071
r: 66153
r: 66204
r: 66135
r: 66139
r: 65686
r: 66084
a: [3,65685,0,2,1,65634]
s: 65685
s: 65702
s: 66072
a: [3,66083,0,6,1,68]
s: 66083
s: 66134
s: 66042
s: 66203
a: [3,65982,0,11,1,98]
s: 65982
s: 66140
s: 66152
d: [2,66252,1,11,2,1,40,65653,65981]
a: [3,66257,1,5,1,66252]
s: 66257
a: [3,66261,1,6,1,66252]
s: 66261
d: [2,66117,1,6,2,0,44,65649,65646]
s: 66117
d: [2,66118,1,6,2,0,44,65649,65646]
s: 66118
d: [2,66372,2,6,2,1,40,66261,65650]
d: [2,66373,2,4,1,2,21,66372,66110]
a: [3,66377,2,4,1,66373]
s: 66377
d: [2,66382,3,0,0,3,52,66203,66134,66377,65642]
e: 66382
d: [2,66384,3,4,2,1,40,66377,65642]
d: [2,66386,3,2,1,2,21,66384,66134]
d: [2,66387,3,0,0,3,21,66386,66203]
e: 66387
r: 65982
r: 65985
r: 66129
r: 66189
r: 66261
r: 66377
r: 65721
r: 66053
r: 66289
s: 65722
s: 65986
s: 66054
s: 65981
s: 66128
a: [3,66290,1,3,1,66286]
s: 66290
d: [2,66400,2,0,0,2,52,65995,66290,65636]
e: 66400
d: [2,66401,2,2,1,1,40,66290,65636]
d: [2,66403,2,0,0,2,21,66401,65995]
e: 66403
r: 65670
r: 65685
r: 65992
r: 65995
r: 66083
r: 66096
r: 66152
r: 66203
r: 65666
r: 65690
r: 66030
r: 66048
r: 66290
a: [3,65665,0,2,1,65634]
s: 65665
s: 65670
s: 65685
a: [3,65689,0,2,1,65634]
s: 65689
s: 65995
s: 66203
s: 65992
s: 66083
s: 66152
a: [3,66029,0,7,1,86]
s: 66029
s: 66189
a: [3,66047,0,7,1,80]
s: 66047
d: [2,66420,1,0,0,2,52,66005,66047,65657]
e: 66420
d: [2,66421,2,0,0,2,52,66128,66047,65655]
e: 66421
d: [2,66422,1,2,1,1,40,66047,65657]
d: [2,66427,1,0,0,2,21,66422,66005]
e: 66427
d: [2,66424,1,3,1,1,40,66047,65655]
d: [2,66429,1,0,0,2,21,66424,66128]
e: 66429
r: 65665
r: 65718
r: 66047
r: 66060
s: 65666
a: [3,65717,0,2,1,65634]
s: 65717
a: [3,66059,0,5,1,76]
s: 66059
d: [2,66431,2,0,0,2,52,66134,66059,65655]
e: 66431
d: [2,66434,1,2,1,1,40,66059,65655]
d: [2,66436,1,0,0,2,21,66434,66134]
e: 66436
r: 65710
r: 65717
r: 66059
r: 66066
a: [3,65709,0,2,1,65634]
s: 65709
s: 65718
s: 66060
s: 66289
a: [3,66065,0,6,1,74]
s: 66065
d: [2,66440,1,0,0,2,52,66002,66065,65657]
e: 66440
d: [2,66441,2,0,0,2,52,66134,66065,65655]
e: 66441
d: [2,66442,1,3,1,1,40,66065,65657]
d: [2,66447,1,0,0,2,21,66442,66002]
e: 66447
d: [2,66444,1,2,1,1,40,66065,65655]
d: [2,66449,1,0,0,2,21,66444,66134]
e: 66449
r: 65678
r: 65709
r: 66065
r: 66090
a: [3,65677,0,2,1,65634]
s: 65677
s: 65710
s: 66066
a: [3,66262,1,6,1,66252]
s: 66262
s: 66048
a: [3,66376,2,4,1,66373]
s: 66376
a: [3,66089,0,7,1,66]
s: 66089
d: [2,66458,1,0,0,2,52,66002,66089,65657]
e: 66458
d: [2,66459,1,3,1,1,40,66089,65657]
d: [2,66463,1,0,0,2,21,66459,66002]
e: 66463
r: 65677
r: 65685
r: 65718
r: 66060
r: 66083
r: 66089
r: 66134
r: 66140
r: 66257
r: 66262
r: 66376
s: 65678
s: 65686
s: 65717
s: 66135
s: 66059
s: 66084
s: 66139
s: 66090
s: 66377
s: 66096
a: [3,66258,1,5,1,66252]
s: 66258
d: [2,66480,2,0,0,2,52,66002,66258,66117]
e: 66480
d: [2,66481,2,3,1,1,40,66258,66117]
d: [2,66484,2,0,0,2,21,66481,66002]
e: 66484
r: 65722
r: 65981
r: 65986
r: 66054
r: 66128
r: 66258
r: 65689
r: 65710
r: 66002
r: 66005
r: 66029
r: 66066
r: 66110
r: 66189
r: 65717
r: 65992
r: 65995
r: 66059
r: 66289
r: 65678
r: 65714
r: 66012
r: 66090
s: 65677
s: 65690
s: 65709
s: 65713
s: 65718
s: 65722
s: 65986
s: 65996
s: 66006
s: 66001
s: 66109
s: 65991
s: 66128
s: 66011
s: 66060
a: [3,66188,1,3,1,66175]
s: 66188
s: 66065
s: 66290
s: 66030
s: 66089
d: [2,66515,2,0,0,2,52,66139,66089,65655]
e: 66515
d: [2,66518,1,3,1,1,40,66089,65655]
d: [2,66520,1,0,0,2,21,66518,66139]
e: 66520
r: 65709
r: 66065
r: 66135
r: 66139
r: 66377
s: 65710
s: 66066
s: 66134
s: 66054
s: 65981
s: 66140
s: 66258
d: [2,66531,2,0,0,2,52,66134,66258,65649]
e: 66531
d: [2,66533,2,2,1,1,40,66258,65649]
d: [2,66535,2,0,0,2,21,66533,66134]
e: 66535
r: 66258
s: 66376
s: 66262
d: [2,66539,2,0,0,2,52,66109,66262,65649]
e: 66539
d: [2,66541,2,2,1,1,40,66262,65649]
d: [2,66543,2,0,0,2,21,66541,66109]
e: 66543
r: 65677
r: 65698
r: 65710
r: 66024
r: 66066
r: 66089
r: 66134
r: 66140
r: 66262
r: 66376
s: 65678
s: 65697
s: 65709
s: 66135
s: 66023
d: [2,66545,2,0,0,2,52,66290,66023,65656]
e: 66545
d: [2,66547,1,3,1,1,40,66023,65656]
d: [2,66550,1,0,0,2,21,66547,66290]
e: 66550
r: 65713
r: 65991
r: 65996
r: 66011
r: 66152
r: 66203
r: 66290
r: 65694
r: 66078
s: 65693
s: 65714
s: 65995
s: 66204
s: 65992
s: 66012
s: 66139
s: 66065
d: [2,66562,2,0,0,2,52,66204,66065,65656]
e: 66562
d: [2,66564,1,2,1,1,40,66065,65656]
d: [2,66567,1,0,0,2,21,66564,66204]
e: 66567
r: 65666
r: 65697
r: 65709
r: 65981
r: 65986
r: 66023
r: 66048
r: 66065
r: 66128
s: 65665
s: 65698
s: 65710
s: 65985
s: 66024
s: 66129
s: 66066
s: 66047
s: 66153
s: 66077
s: 66289
s: 66090
s: 65982
s: 66377
d: [2,66586,3,0,0,3,52,66204,66135,66377,65643]
e: 66586
d: [2,66587,3,4,2,1,40,66377,65643]
d: [2,66590,3,2,1,2,21,66587,66135]
d: [2,66591,3,0,0,3,21,66590,66204]
e: 66591
r: 66377
s: 66257
d: [2,66593,2,0,0,3,52,66135,66001,66257,65650]
e: 66593
d: [2,66594,2,5,2,1,40,66257,65650]
d: [2,66596,2,2,1,2,21,66594,66001]
d: [2,66597,2,0,0,3,21,66596,66135]
e: 66597
r: 65662
r: 65682
r: 65693
r: 66036
r: 66077
r: 66102
r: 66135
r: 66139
s: 65661
s: 65681
s: 65694
s: 66134
s: 66035
d: [2,66600,1,0,0,2,52,65995,66035,65656]
e: 66600
d: [2,66602,1,2,1,1,40,66035,65656]
d: [2,66605,1,0,0,2,21,66602,65995]
e: 66605
r: 65678
r: 65681
r: 66035
r: 66090
s: 65677
s: 65682
s: 66036
s: 66140
s: 66078
s: 66261
s: 66089
d: [2,66611,2,0,0,2,52,66204,66089,65656]
e: 66611
d: [2,66613,1,2,1,1,40,66089,65656]
d: [2,66616,1,0,0,2,21,66613,66204]
e: 66616
r: 65677
r: 65702
r: 66072
r: 66089
s: 65678
s: 65701
s: 66071
d: [2,66618,2,0,0,2,52,66109,66071,65657]
e: 66618
d: [2,66619,2,0,0,2,52,66134,66071,65655]
e: 66619
d: [2,66620,1,2,1,1,40,66071,65657]
d: [2,66625,1,0,0,2,21,66620,66109]
e: 66625
d: [2,66622,1,2,1,1,40,66071,65655]
d: [2,66627,1,0,0,2,21,66622,66134]
e: 66627
r: 65661
r: 65665
r: 65682
r: 65701
r: 65992
r: 65995
r: 66036
r: 66047
r: 66071
r: 66153
r: 66204
r: 66289
s: 65662
s: 65666
s: 65681
s: 65702
s: 65996
s: 66203
s: 65991
s: 66035
s: 66152
s: 66072
s: 66290
s: 66048
s: 66090
d: [2,66638,2,0,0,4,52,66203,66140,66001,66090,65658]
e: 66638
d: [2,66639,1,8,3,1,40,66090,65658]
d: [2,66641,1,5,2,2,21,66639,66001]
d: [2,66642,1,3,1,3,21,66641,66203]
d: [2,66643,1,0,0,4,21,66642,66140]
e: 66643
r: 65678
r: 65686
r: 66084
r: 66090
s: 65677
s: 65685
s: 66083
d: [2,66645,2,0,0,2,52,66109,66083,65657]
e: 66645
d: [2,66646,1,2,1,1,40,66083,65657]
d: [2,66650,1,0,0,2,21,66646,66109]
e: 66650
r: 65674
r: 65685
r: 66042
r: 66083
s: 65673
s: 65686
s: 66084
s: 66041
d: [2,66653,2,0,0,2,52,66188,66041,65657]
e: 66653
d: [2,66654,1,3,1,1,40,66041,65657]
d: [2,66658,1,0,0,2,21,66654,66188]
e: 66658
r: 65670
r: 65677
r: 66001
r: 66006
r: 66096
r: 66109
r: 66152
r: 66188
r: 66203
r: 65673
r: 65681
r: 66035
r: 66041
s: 65669
s: 65674
s: 65678
s: 65682
s: 66005
s: 66110
s: 66002
s: 66036
s: 66204
s: 66042
s: 66153
s: 66090
s: 66095
s: 66188
d: [2,66678,2,0,0,2,52,66005,66188,65636]
e: 66678
d: [2,66679,2,2,1,1,40,66188,65636]
d: [2,66681,2,0,0,2,21,66679,66005]
e: 66681
r: 65669
r: 65694
r: 65702
r: 65706
r: 65982
r: 65985
r: 66018
r: 66072
r: 66078
r: 66095
r: 66129
r: 66134
r: 66140
r: 66188
r: 66261
s: 65670
s: 65693
s: 65701
s: 65705
s: 65986
s: 66017
s: 66135
s: 66071
s: 66077
d: [2,66691,1,0,0,2,52,66002,66077,65657]
e: 66691
d: [2,66692,1,3,1,1,40,66077,65657]
d: [2,66696,1,0,0,2,21,66692,66002]
e: 66696
r: 65693
r: 65718
r: 66060
r: 66077
s: 65694
s: 65717
s: 66128
s: 66059
d: [2,66699,2,0,0,2,52,66204,66059,65656]
e: 66699
d: [2,66701,1,2,1,1,40,66059,65656]
d: [2,66704,1,0,0,2,21,66701,66204]
e: 66704
r: 65670
r: 65701
r: 65717
r: 66059
r: 66071
r: 66135
s: 65669
s: 65702
s: 65718
s: 66134
s: 66060
s: 66072
s: 66140
s: 66078
s: 66189
s: 66095
s: 66102
s: 66376
d: [2,66715,3,0,0,3,52,66134,66204,66376,65640]
e: 66715
d: [2,66716,3,4,2,1,40,66376,65640]
d: [2,66719,3,2,1,2,21,66716,66204]
d: [2,66720,3,0,0,3,21,66719,66134]
e: 66720
i: [1,49251,0,32,16,1,0,0] t1_gate_2
i: [1,12387,0,32,16,1,0,0] t1_gate_2
i: [1,3171,0,32,16,1,0,0] t1_gate_2
i: [1,867,0,32,16,1,0,0] t1_gate_2
i: [1,291,0,32,16,1,0,0] t1_gate_2
i: [1,147,0,32,16,1,0,0] t1_gate_2
i: [1,111,0,32,16,1,0,0] t1_gate_2
i: [1,102,0,32,16,1,0,0] t1_gate_2
i: [1,97,0,13,2,1,0,0] t1_gate_2
i: [1,95,0,6,2,1,0,0] t1_gate_2
i: [1,93,0,5,2,1,0,0] t1_gate_2
i: [1,91,0,7,2,1,0,0] t1_gate_2
i: [1,89,0,8,2,1,0,0] t1_gate_2
i: [1,87,0,8,2,1,0,0] t1_gate_2
i: [1,85,0,9,2,1,0,0] t1_gate_2
i: [1,83,0,8,2,1,0,0] t1_gate_2
i: [1,81,0,9,2,1,0,0] t1_gate_2
i: [1,79,0,9,2,1,0,0] t1_gate_2
i: [1,77,0,10,2,1,0,0] t1_gate_2
i: [1,75,0,7,2,1,0,0] t1_gate_2
i: [1,73,0,8,2,1,0,0] t1_gate_2
i: [1,71,0,8,2,1,0,0] t1_gate_2
i: [1,69,0,9,2,1,0,0] t1_gate_2
i: [1,67,0,8,2,1,0,0] t1_gate_2
i: [1,65,0,9,2,1,0,0] t1_gate_2
i: [1,63,0,9,2,1,0,0] t1_gate_2
i: [1,61,0,10,2,1,0,0] t1_gate_2
d: [2,66122,1,3,1,1,40,65655,66041]
d: [2,66125,1,3,1,1,40,65655,66095]
d: [2,66148,1,3,1,1,40,65656,66101]
d: [2,66149,1,3,1,1,40,65656,66095]
d: [2,66120,1,2,1,1,40,65655,66017]
d: [2,66119,1,2,1,1,40,65655,66011]
d: [2,66177,1,2,1,1,40,65657,66095]
d: [2,66144,1,2,1,1,40,65656,66017]
d: [2,66143,1,2,1,1,40,65656,66011]
d: [2,66238,2,2,1,1,40,66152,65636]
d: [2,66248,1,2,1,1,40,66002,65636]
d: [2,66165,2,2,1,1,40,66128,65636]
d: [2,66271,1,8,3,1,40,65658,66096]
d: [2,66282,1,6,2,2,21,66271,66110]
d: [2,66323,1,3,1,1,40,66071,65656]
d: [2,66331,1,3,1,1,40,66053,65656]
d: [2,66327,1,2,1,1,40,66041,65656]
d: [2,66348,1,2,1,2,21,66343,66110]
d: [2,66275,1,6,3,1,40,65658,66060]
d: [2,66350,1,4,2,2,21,66275,66110]
d: [2,66356,1,7,3,1,40,66072,65658]
d: [2,66357,1,5,2,2,21,66356,66110]
d: [2,66360,1,2,1,1,40,66083,65656]
d: [2,66361,1,3,1,1,40,66083,65655]
d: [2,66332,1,3,1,1,40,66053,65655]
d: [2,66330,1,3,1,1,40,66053,65657]
d: [2,66394,1,9,3,1,40,66054,65658]
d: [2,66414,1,3,1,1,40,66029,65657]
d: [2,66415,1,3,1,1,40,66029,65656]
d: [2,66452,2,4,1,1,40,66262,66117]
d: [2,66457,3,4,2,1,40,66376,65639]
d: [2,66465,3,2,1,2,21,66457,66203]
d: [2,66416,1,2,1,1,40,66029,65655]
d: [2,66468,1,2,1,1,40,66059,65657]
d: [2,66408,1,4,2,1,40,65992,65643]
d: [2,66492,1,2,1,2,21,66408,66203]
d: [2,66273,1,7,3,1,40,65658,66036]
d: [2,66281,1,6,3,1,40,65658,66012]
d: [2,66301,1,7,3,1,40,66018,65658]
d: [2,66321,1,7,3,1,40,66024,65658]
d: [2,66455,1,8,3,1,40,66048,65658]
d: [2,66500,1,2,1,1,40,66001,65635]
d: [2,66523,1,7,3,1,40,66066,65658]
d: [2,66524,1,4,2,2,21,66523,66001]
d: [2,66525,1,2,1,3,21,66524,66203]
d: [2,66528,1,11,2,1,40,65981,65653]
d: [2,66272,1,9,3,1,40,65658,66102]
d: [2,66552,1,6,2,2,21,66272,66001]
d: [2,66278,1,8,3,1,40,65658,66078]
d: [2,66554,1,5,2,2,21,66278,66001]
d: [2,66555,1,3,1,3,21,66554,66135]
d: [2,66573,1,3,1,1,40,66047,65656]
d: [2,66582,2,5,1,1,40,65982,66113]
d: [2,66583,1,6,1,1,40,65982,65648]
d: [2,66631,1,2,1,1,40,66035,65657]
d: [2,66652,1,7,3,1,40,66084,65658]
d: [2,66661,1,4,2,2,21,66652,66140]
d: [2,66630,1,4,2,1,40,65991,65639]
d: [2,66610,2,6,2,1,40,66261,65650]
d: [2,66629,1,4,2,1,40,65991,65640]
d: [2,66669,1,8,3,1,40,66042,65658]
d: [2,66670,1,6,2,2,21,66669,65996]
d: [2,66684,1,3,1,1,40,66017,65657]
f: 66720,66704,66696,66684,66670,66681,66629,66654,66610,66630,66661,66631,66650,66643,66627,66625,66616,66605,66596,66591,66583,66582,66573,66567,66547,66555,66552,66543,66535,66528,66525,66518,66500,66455,66321,66301,66281,66273,66492,66468,66416,66484,66465,66463,66452,66449,66447,66436,66429,66427,66415,66414,66401,66394,66280,66330,66332,66387,66361,66360,66357,66350,66348,66327,66335,66331,66323,66319,66317,66303,66282,66165,66248,66238,66240,66224,66143,66144,66177,66174,66176,66167,66119,66120,66158,66149,66148,66147,66125,66124,66123,66122,66121,61,62,63,64,65,66,67,68,69,70,71,72,73,74,75,76,77,78,79,80,81,82,83,84,85,86,87,88,89,90,91,92,93,94,95,96,97,98,102,111,147,291,867,3171,12387,49251
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_2__t1_gate_2
% SZS output start Proof for gate_2__t1_gate_2
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ~(~(~(v1_xboole_0(X8) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18)))) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k4_gate_1(X16,X17))) & ~(v1_xboole_0(k4_gate_1(X16,X17)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k1_gate_1(X16))) & ~(v1_xboole_0(k1_gate_1(X16)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [input t1_gate_2]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ! [X19] : ! [X20] : ! [X21] : ~(~(~(v1_xboole_0(X8) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18)))) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k4_gate_1(X16,X17))) & ~(v1_xboole_0(k4_gate_1(X16,X17)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k1_gate_1(X16))) & ~(v1_xboole_0(k1_gate_1(X16)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [negated conjecture 1]
4. ! [X0,X1] : k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [input commutativity_k2_gate_1]
5. ! [X0,X1] : k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [input commutativity_k3_gate_1]
13. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc17_gate_1]
14. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc18_gate_1]
15. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X0,X2))) [input fc19_gate_1]
16. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
17. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X2,X0))) [input fc20_gate_1]
18. ! [X0] : (~v1_xboole_0(X0) => v1_xboole_0(k1_gate_1(X0))) [input fc2_gate_1]
19. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_gate_1(X0,X1))) [input fc3_gate_1]
20. ! [X0,X1] : (v1_xboole_0(X0) => v1_xboole_0(k2_gate_1(X0,X1))) [input fc4_gate_1]
21. ! [X0,X1] : (~v1_xboole_0(X1) => ~v1_xboole_0(k3_gate_1(X0,X1))) [input fc5_gate_1]
22. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k3_gate_1(X0,X1))) [input fc6_gate_1]
27. ! [X0] : ! [X1] : (~(v1_xboole_0(k4_gate_1(X0,X1)) & ((~v1_xboole_0(X1) & v1_xboole_0(X0)) | (v1_xboole_0(X1) & ~v1_xboole_0(X0)))) & ~(~(~v1_xboole_0(X1) & v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~v1_xboole_0(k4_gate_1(X0,X1)))) [input t6_gate_1]
30. ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : ~(~(~(v1_xboole_0(X8) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(X0)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X9))) & ~(v1_xboole_0(X21) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18)))) & ~v1_xboole_0(X21)) & ~(v1_xboole_0(X20) & ~v1_xboole_0(k4_gate_1(X16,X17))) & ~(v1_xboole_0(k4_gate_1(X16,X17)) & ~v1_xboole_0(X20)) & ~(v1_xboole_0(X19) & ~v1_xboole_0(k1_gate_1(X16))) & ~(v1_xboole_0(k1_gate_1(X16)) & ~v1_xboole_0(X19)) & ~(v1_xboole_0(X15) & ~v1_xboole_0(k8_gate_1(X21,X20,X19))) & ~(v1_xboole_0(k8_gate_1(X21,X20,X19)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(X14) & ~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(X13) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X12)) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) & ~v1_xboole_0(X11)) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) & ~v1_xboole_0(X10)) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) & ~v1_xboole_0(X9)) & ~(v1_xboole_0(X8) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) & ~v1_xboole_0(X8)) & ~(v1_xboole_0(X7) & ~v1_xboole_0(k8_gate_1(X18,X17,X16))) & ~(v1_xboole_0(k8_gate_1(X18,X17,X16)) & ~v1_xboole_0(X7)) & ~(v1_xboole_0(X6) & ~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) & ~v1_xboole_0(X6)) & ~(v1_xboole_0(X5) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) & ~v1_xboole_0(X5)) & ~(v1_xboole_0(X4) & ~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X4)) & ~(v1_xboole_0(X3) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) & ~v1_xboole_0(X3)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) & ~v1_xboole_0(X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & ~(v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) & ~v1_xboole_0(X0))) [flattening 2]
31. ! [X0,X1] : (~(v1_xboole_0(k4_gate_1(X0,X1)) & ((~v1_xboole_0(X1) & v1_xboole_0(X0)) | (v1_xboole_0(X1) & ~v1_xboole_0(X0)))) & ~(~(~v1_xboole_0(X1) & v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~v1_xboole_0(k4_gate_1(X0,X1)))) [flattening 27]
36. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (((v1_xboole_0(X8) & ~v1_xboole_0(X7)) | (v1_xboole_0(X7) & ~v1_xboole_0(X8)) | (v1_xboole_0(X15) & ~v1_xboole_0(X6)) | (v1_xboole_0(X6) & ~v1_xboole_0(X15)) | (v1_xboole_0(X14) & ~v1_xboole_0(X5)) | (v1_xboole_0(X5) & ~v1_xboole_0(X14)) | (v1_xboole_0(X13) & ~v1_xboole_0(X4)) | (v1_xboole_0(X4) & ~v1_xboole_0(X13)) | (v1_xboole_0(X12) & ~v1_xboole_0(X3)) | (v1_xboole_0(X3) & ~v1_xboole_0(X12)) | (v1_xboole_0(X11) & ~v1_xboole_0(X2)) | (v1_xboole_0(X2) & ~v1_xboole_0(X11)) | (v1_xboole_0(X10) & ~v1_xboole_0(X1)) | (v1_xboole_0(X1) & ~v1_xboole_0(X10)) | (v1_xboole_0(X9) & ~v1_xboole_0(X0)) | (v1_xboole_0(X0) & ~v1_xboole_0(X9))) & (~v1_xboole_0(X21) | v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))))) & (~v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18)))) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(k4_gate_1(X16,X17))) & (~v1_xboole_0(k4_gate_1(X16,X17)) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k1_gate_1(X16))) & (~v1_xboole_0(k1_gate_1(X16)) | v1_xboole_0(X19)) & (~v1_xboole_0(X15) | v1_xboole_0(k8_gate_1(X21,X20,X19))) & (~v1_xboole_0(k8_gate_1(X21,X20,X19)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) | v1_xboole_0(X10)) & (~v1_xboole_0(X9) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) | v1_xboole_0(X9)) & (~v1_xboole_0(X8) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X8)) & (~v1_xboole_0(X7) | v1_xboole_0(k8_gate_1(X18,X17,X16))) & (~v1_xboole_0(k8_gate_1(X18,X17,X16)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X0))) [ennf transformation 30]
37. ! [X0] : (v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0)) [ennf transformation 18]
38. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 16]
40. ! [X0,X1] : ((~v1_xboole_0(k4_gate_1(X0,X1)) | ((v1_xboole_0(X1) | ~v1_xboole_0(X0)) & (~v1_xboole_0(X1) | v1_xboole_0(X0)))) & ((~v1_xboole_0(X1) & v1_xboole_0(X0)) | (v1_xboole_0(X1) & ~v1_xboole_0(X0)) | v1_xboole_0(k4_gate_1(X0,X1)))) [ennf transformation 31]
42. ! [X0,X1] : (~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1)) [ennf transformation 21]
43. ! [X0,X1] : (v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 20]
44. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 19]
45. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 44]
50. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 22]
51. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 50]
54. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0)) [ennf transformation 14]
55. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0)) [ennf transformation 15]
56. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0)) [ennf transformation 17]
57. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 13]
58. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 57]
59. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (((v1_xboole_0(X8) & ~v1_xboole_0(X7)) | (v1_xboole_0(X7) & ~v1_xboole_0(X8)) | (v1_xboole_0(X15) & ~v1_xboole_0(X6)) | (v1_xboole_0(X6) & ~v1_xboole_0(X15)) | (v1_xboole_0(X14) & ~v1_xboole_0(X5)) | (v1_xboole_0(X5) & ~v1_xboole_0(X14)) | (v1_xboole_0(X13) & ~v1_xboole_0(X4)) | (v1_xboole_0(X4) & ~v1_xboole_0(X13)) | (v1_xboole_0(X12) & ~v1_xboole_0(X3)) | (v1_xboole_0(X3) & ~v1_xboole_0(X12)) | (v1_xboole_0(X11) & ~v1_xboole_0(X2)) | (v1_xboole_0(X2) & ~v1_xboole_0(X11)) | (v1_xboole_0(X10) & ~v1_xboole_0(X1)) | (v1_xboole_0(X1) & ~v1_xboole_0(X10)) | (v1_xboole_0(X9) & ~v1_xboole_0(X0)) | (v1_xboole_0(X0) & ~v1_xboole_0(X9))) & (~v1_xboole_0(X21) | v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18))))) & (~v1_xboole_0(k3_gate_1(k2_gate_1(X18,k1_gate_1(X16)),k2_gate_1(X16,k4_gate_1(X17,X18)))) | v1_xboole_0(X21)) & (~v1_xboole_0(X20) | v1_xboole_0(k4_gate_1(X16,X17))) & (~v1_xboole_0(k4_gate_1(X16,X17)) | v1_xboole_0(X20)) & (~v1_xboole_0(X19) | v1_xboole_0(k1_gate_1(X16))) & (~v1_xboole_0(k1_gate_1(X16)) | v1_xboole_0(X19)) & (~v1_xboole_0(X15) | v1_xboole_0(k8_gate_1(X21,X20,X19))) & (~v1_xboole_0(k8_gate_1(X21,X20,X19)) | v1_xboole_0(X15)) & (~v1_xboole_0(X14) | v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,X20,k1_gate_1(X19))) | v1_xboole_0(X14)) & (~v1_xboole_0(X13) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),X19)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(X21,k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X12)) & (~v1_xboole_0(X11) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,X19)) | v1_xboole_0(X11)) & (~v1_xboole_0(X10) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),X20,k1_gate_1(X19))) | v1_xboole_0(X10)) & (~v1_xboole_0(X9) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),X19)) | v1_xboole_0(X9)) & (~v1_xboole_0(X8) | v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X21),k1_gate_1(X20),k1_gate_1(X19))) | v1_xboole_0(X8)) & (~v1_xboole_0(X7) | v1_xboole_0(k8_gate_1(X18,X17,X16))) & (~v1_xboole_0(k8_gate_1(X18,X17,X16)) | v1_xboole_0(X7)) & (~v1_xboole_0(X6) | v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,X17,k1_gate_1(X16))) | v1_xboole_0(X6)) & (~v1_xboole_0(X5) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),X16)) | v1_xboole_0(X5)) & (~v1_xboole_0(X4) | v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(X18,k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X4)) & (~v1_xboole_0(X3) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,X16)) | v1_xboole_0(X3)) & (~v1_xboole_0(X2) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),X17,k1_gate_1(X16))) | v1_xboole_0(X2)) & (~v1_xboole_0(X1) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),X16)) | v1_xboole_0(X1)) & (~v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(X18),k1_gate_1(X17),k1_gate_1(X16))) | v1_xboole_0(X0))) => (((v1_xboole_0(sK8) & ~v1_xboole_0(sK7)) | (v1_xboole_0(sK7) & ~v1_xboole_0(sK8)) | (v1_xboole_0(sK15) & ~v1_xboole_0(sK6)) | (v1_xboole_0(sK6) & ~v1_xboole_0(sK15)) | (v1_xboole_0(sK14) & ~v1_xboole_0(sK5)) | (v1_xboole_0(sK5) & ~v1_xboole_0(sK14)) | (v1_xboole_0(sK13) & ~v1_xboole_0(sK4)) | (v1_xboole_0(sK4) & ~v1_xboole_0(sK13)) | (v1_xboole_0(sK12) & ~v1_xboole_0(sK3)) | (v1_xboole_0(sK3) & ~v1_xboole_0(sK12)) | (v1_xboole_0(sK11) & ~v1_xboole_0(sK2)) | (v1_xboole_0(sK2) & ~v1_xboole_0(sK11)) | (v1_xboole_0(sK10) & ~v1_xboole_0(sK1)) | (v1_xboole_0(sK1) & ~v1_xboole_0(sK10)) | (v1_xboole_0(sK9) & ~v1_xboole_0(sK0)) | (v1_xboole_0(sK0) & ~v1_xboole_0(sK9))) & (~v1_xboole_0(sK21) | v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))))) & (~v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(k4_gate_1(sK16,sK17))) & (~v1_xboole_0(k4_gate_1(sK16,sK17)) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK16))) & (~v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0))) [choice axiom]
60. ((v1_xboole_0(sK8) & ~v1_xboole_0(sK7)) | (v1_xboole_0(sK7) & ~v1_xboole_0(sK8)) | (v1_xboole_0(sK15) & ~v1_xboole_0(sK6)) | (v1_xboole_0(sK6) & ~v1_xboole_0(sK15)) | (v1_xboole_0(sK14) & ~v1_xboole_0(sK5)) | (v1_xboole_0(sK5) & ~v1_xboole_0(sK14)) | (v1_xboole_0(sK13) & ~v1_xboole_0(sK4)) | (v1_xboole_0(sK4) & ~v1_xboole_0(sK13)) | (v1_xboole_0(sK12) & ~v1_xboole_0(sK3)) | (v1_xboole_0(sK3) & ~v1_xboole_0(sK12)) | (v1_xboole_0(sK11) & ~v1_xboole_0(sK2)) | (v1_xboole_0(sK2) & ~v1_xboole_0(sK11)) | (v1_xboole_0(sK10) & ~v1_xboole_0(sK1)) | (v1_xboole_0(sK1) & ~v1_xboole_0(sK10)) | (v1_xboole_0(sK9) & ~v1_xboole_0(sK0)) | (v1_xboole_0(sK0) & ~v1_xboole_0(sK9))) & (~v1_xboole_0(sK21) | v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18))))) & (~v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) | v1_xboole_0(sK21)) & (~v1_xboole_0(sK20) | v1_xboole_0(k4_gate_1(sK16,sK17))) & (~v1_xboole_0(k4_gate_1(sK16,sK17)) | v1_xboole_0(sK20)) & (~v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK16))) & (~v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK19)) & (~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15)) & (~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14)) & (~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12)) & (~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11)) & (~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10)) & (~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9)) & (~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8)) & (~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7)) & (~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6)) & (~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5)) & (~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4)) & (~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3)) & (~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16)))) & (~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0)) [skolemisation 36,59]
61. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK0) [cnf transformation 60] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
62. ~v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) [cnf transformation 60] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
63. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) | v1_xboole_0(sK1) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
64. ~v1_xboole_0(sK1) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) | v1_xboole_0(sK2) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
66. ~v1_xboole_0(sK2) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
67. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) | v1_xboole_0(sK3) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
68. ~v1_xboole_0(sK3) | v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) | v1_xboole_0(sK4) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. ~v1_xboole_0(sK4) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
71. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) | v1_xboole_0(sK5) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
72. ~v1_xboole_0(sK5) | v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
73. ~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) | v1_xboole_0(sK6) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
74. ~v1_xboole_0(sK6) | v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
75. ~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) | v1_xboole_0(sK7) [cnf transformation 60] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
76. ~v1_xboole_0(sK7) | v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) [cnf transformation 60] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
77. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK8) [cnf transformation 60] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
78. ~v1_xboole_0(sK8) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) [cnf transformation 60] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
79. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) | v1_xboole_0(sK9) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
80. ~v1_xboole_0(sK9) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
81. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) | v1_xboole_0(sK10) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
82. ~v1_xboole_0(sK10) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
83. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) | v1_xboole_0(sK11) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
84. ~v1_xboole_0(sK11) | v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
85. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) | v1_xboole_0(sK12) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
86. ~v1_xboole_0(sK12) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) [cnf transformation 60] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
87. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) | v1_xboole_0(sK13) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
88. ~v1_xboole_0(sK13) | v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
89. ~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) | v1_xboole_0(sK14) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
90. ~v1_xboole_0(sK14) | v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) [cnf transformation 60] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
91. ~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) | v1_xboole_0(sK15) [cnf transformation 60] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
92. ~v1_xboole_0(sK15) | v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) [cnf transformation 60] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
93. ~v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK19) [cnf transformation 60] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
94. ~v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK16)) [cnf transformation 60] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
95. ~v1_xboole_0(k4_gate_1(sK16,sK17)) | v1_xboole_0(sK20) [cnf transformation 60] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
96. ~v1_xboole_0(sK20) | v1_xboole_0(k4_gate_1(sK16,sK17)) [cnf transformation 60] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
97. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) | v1_xboole_0(sK21) [cnf transformation 60] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
98. ~v1_xboole_0(sK21) | v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) [cnf transformation 60] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
102. ~v1_xboole_0(sK7) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK10) | v1_xboole_0(sK9) | v1_xboole_0(sK0) [cnf transformation 60] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
111. ~v1_xboole_0(sK7) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK11) | v1_xboole_0(sK10) | v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 60] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
147. ~v1_xboole_0(sK7) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK12) | v1_xboole_0(sK11) | v1_xboole_0(sK2) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 60] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
291. ~v1_xboole_0(sK7) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK13) | v1_xboole_0(sK12) | v1_xboole_0(sK3) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 60] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
867. ~v1_xboole_0(sK7) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK14) | v1_xboole_0(sK13) | v1_xboole_0(sK4) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 60] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
3171. ~v1_xboole_0(sK7) | ~v1_xboole_0(sK8) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK15) | v1_xboole_0(sK14) | v1_xboole_0(sK5) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 60] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
12387. ~v1_xboole_0(sK7) | ~v1_xboole_0(sK8) | v1_xboole_0(sK15) | v1_xboole_0(sK6) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 60] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49251. v1_xboole_0(sK8) | v1_xboole_0(sK7) | ~v1_xboole_0(sK6) | ~v1_xboole_0(sK15) | ~v1_xboole_0(sK5) | ~v1_xboole_0(sK14) | ~v1_xboole_0(sK4) | ~v1_xboole_0(sK13) | ~v1_xboole_0(sK3) | ~v1_xboole_0(sK12) | ~v1_xboole_0(sK2) | ~v1_xboole_0(sK11) | ~v1_xboole_0(sK1) | ~v1_xboole_0(sK10) | ~v1_xboole_0(sK0) | ~v1_xboole_0(sK9) [cnf transformation 60] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65635. v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0) [cnf transformation 37] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65636. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 38] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65639. v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65640. v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X0) | ~v1_xboole_0(X1) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65642. ~v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65643. ~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65645. k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [cnf transformation 5] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65646. k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65648. ~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1) [cnf transformation 42] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65649. v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 43] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65650. ~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 45] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65653. v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 51] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65655. v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0) [cnf transformation 54] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65656. v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0) [cnf transformation 55] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65657. v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0) [cnf transformation 56] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65658. ~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 58] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65660. 1 <=> v1_xboole_0(sK0) [avatar definition]
65664. 2 <=> v1_xboole_0(sK9) [avatar definition]
65668. 3 <=> v1_xboole_0(sK1) [avatar definition]
65672. 4 <=> v1_xboole_0(sK10) [avatar definition]
65676. 5 <=> v1_xboole_0(sK2) [avatar definition]
65680. 6 <=> v1_xboole_0(sK11) [avatar definition]
65684. 7 <=> v1_xboole_0(sK3) [avatar definition]
65688. 8 <=> v1_xboole_0(sK12) [avatar definition]
65692. 9 <=> v1_xboole_0(sK4) [avatar definition]
65696. 10 <=> v1_xboole_0(sK13) [avatar definition]
65700. 11 <=> v1_xboole_0(sK5) [avatar definition]
65704. 12 <=> v1_xboole_0(sK14) [avatar definition]
65708. 13 <=> v1_xboole_0(sK6) [avatar definition]
65712. 14 <=> v1_xboole_0(sK15) [avatar definition]
65716. 15 <=> v1_xboole_0(sK7) [avatar definition]
65720. 16 <=> v1_xboole_0(sK8) [avatar definition]
65850. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | 15 | 16 [avatar split clause 49251,65720,65716,65708,65712,65700,65704,65692,65696,65684,65688,65676,65680,65668,65672,65660,65664]
65914. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | 13 | 14 | ~16 | ~15 [avatar split clause 12387,65716,65720,65712,65708,65700,65704,65692,65696,65684,65688,65676,65680,65668,65672,65660,65664]
65946. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | 11 | 12 | ~14 | ~13 | ~16 | ~15 [avatar split clause 3171,65716,65720,65708,65712,65704,65700,65692,65696,65684,65688,65676,65680,65668,65672,65660,65664]
65962. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | 9 | 10 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 867,65716,65720,65708,65712,65700,65704,65696,65692,65684,65688,65676,65680,65668,65672,65660,65664]
65970. ~2 | ~1 | ~4 | ~3 | ~6 | ~5 | 7 | 8 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 291,65716,65720,65708,65712,65700,65704,65692,65696,65688,65684,65676,65680,65668,65672,65660,65664]
65974. ~2 | ~1 | ~4 | ~3 | 5 | 6 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 147,65716,65720,65708,65712,65700,65704,65692,65696,65684,65688,65680,65676,65668,65672,65660,65664]
65976. ~2 | ~1 | 3 | 4 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 111,65716,65720,65708,65712,65700,65704,65692,65696,65684,65688,65676,65680,65672,65668,65660,65664]
65977. 1 | 2 | ~4 | ~3 | ~6 | ~5 | ~8 | ~7 | ~10 | ~9 | ~12 | ~11 | ~14 | ~13 | ~16 | ~15 [avatar split clause 102,65716,65720,65708,65712,65700,65704,65692,65696,65684,65688,65676,65680,65668,65672,65664,65660]
65980. 17 <=> v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) [avatar definition]
65981. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) <- (~17) [avatar component clause 65980] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65982. v1_xboole_0(k3_gate_1(k2_gate_1(sK18,k1_gate_1(sK16)),k2_gate_1(sK16,k4_gate_1(sK17,sK18)))) <- (17) [avatar component clause 65980] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65984. 18 <=> v1_xboole_0(sK21) [avatar definition]
65985. v1_xboole_0(sK21) <- (18) [avatar component clause 65984] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65987. 17 | ~18 [avatar split clause 98,65984,65980]
65988. 18 | ~17 [avatar split clause 97,65980,65984]
65990. 19 <=> v1_xboole_0(k4_gate_1(sK16,sK17)) [avatar definition]
65991. ~v1_xboole_0(k4_gate_1(sK16,sK17)) <- (~19) [avatar component clause 65990] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65992. v1_xboole_0(k4_gate_1(sK16,sK17)) <- (19) [avatar component clause 65990] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65994. 20 <=> v1_xboole_0(sK20) [avatar definition]
65995. v1_xboole_0(sK20) <- (20) [avatar component clause 65994] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65996. ~v1_xboole_0(sK20) <- (~20) [avatar component clause 65994] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65997. 19 | ~20 [avatar split clause 96,65994,65990]
65998. 20 | ~19 [avatar split clause 95,65990,65994]
66000. 21 <=> v1_xboole_0(k1_gate_1(sK16)) [avatar definition]
66001. ~v1_xboole_0(k1_gate_1(sK16)) <- (~21) [avatar component clause 66000] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66002. v1_xboole_0(k1_gate_1(sK16)) <- (21) [avatar component clause 66000] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66004. 22 <=> v1_xboole_0(sK19) [avatar definition]
66005. v1_xboole_0(sK19) <- (22) [avatar component clause 66004] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66007. 21 | ~22 [avatar split clause 94,66004,66000]
66008. 22 | ~21 [avatar split clause 93,66000,66004]
66010. 23 <=> v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) [avatar definition]
66011. ~v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) <- (~23) [avatar component clause 66010] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66012. v1_xboole_0(k8_gate_1(sK21,sK20,sK19)) <- (23) [avatar component clause 66010] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66013. 23 | ~14 [avatar split clause 92,65712,66010]
66014. 14 | ~23 [avatar split clause 91,66010,65712]
66016. 24 <=> v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) [avatar definition]
66017. ~v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) <- (~24) [avatar component clause 66016] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66018. v1_xboole_0(k8_gate_1(sK21,sK20,k1_gate_1(sK19))) <- (24) [avatar component clause 66016] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66019. 24 | ~12 [avatar split clause 90,65704,66016]
66020. 12 | ~24 [avatar split clause 89,66016,65704]
66022. 25 <=> v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) [avatar definition]
66023. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) <- (~25) [avatar component clause 66022] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66024. v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),sK19)) <- (25) [avatar component clause 66022] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66025. 25 | ~10 [avatar split clause 88,65696,66022]
66026. 10 | ~25 [avatar split clause 87,66022,65696]
66028. 26 <=> v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) [avatar definition]
66029. ~v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) <- (~26) [avatar component clause 66028] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66030. v1_xboole_0(k8_gate_1(sK21,k1_gate_1(sK20),k1_gate_1(sK19))) <- (26) [avatar component clause 66028] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66031. 26 | ~8 [avatar split clause 86,65688,66028]
66032. 8 | ~26 [avatar split clause 85,66028,65688]
66034. 27 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) [avatar definition]
66035. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) <- (~27) [avatar component clause 66034] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66036. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,sK19)) <- (27) [avatar component clause 66034] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66037. 27 | ~6 [avatar split clause 84,65680,66034]
66038. 6 | ~27 [avatar split clause 83,66034,65680]
66040. 28 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) [avatar definition]
66041. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) <- (~28) [avatar component clause 66040] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66042. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),sK20,k1_gate_1(sK19))) <- (28) [avatar component clause 66040] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66043. 28 | ~4 [avatar split clause 82,65672,66040]
66044. 4 | ~28 [avatar split clause 81,66040,65672]
66046. 29 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) [avatar definition]
66047. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) <- (~29) [avatar component clause 66046] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66048. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),sK19)) <- (29) [avatar component clause 66046] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66049. 29 | ~2 [avatar split clause 80,65664,66046]
66050. 2 | ~29 [avatar split clause 79,66046,65664]
66052. 30 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) [avatar definition]
66053. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) <- (~30) [avatar component clause 66052] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66054. v1_xboole_0(k8_gate_1(k1_gate_1(sK21),k1_gate_1(sK20),k1_gate_1(sK19))) <- (30) [avatar component clause 66052] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66055. 30 | ~16 [avatar split clause 78,65720,66052]
66056. 16 | ~30 [avatar split clause 77,66052,65720]
66058. 31 <=> v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) [avatar definition]
66059. ~v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) <- (~31) [avatar component clause 66058] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66060. v1_xboole_0(k8_gate_1(sK18,sK17,sK16)) <- (31) [avatar component clause 66058] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66061. 31 | ~15 [avatar split clause 76,65716,66058]
66062. 15 | ~31 [avatar split clause 75,66058,65716]
66064. 32 <=> v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) [avatar definition]
66065. ~v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) <- (~32) [avatar component clause 66064] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66066. v1_xboole_0(k8_gate_1(sK18,sK17,k1_gate_1(sK16))) <- (32) [avatar component clause 66064] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66067. 32 | ~13 [avatar split clause 74,65708,66064]
66068. 13 | ~32 [avatar split clause 73,66064,65708]
66070. 33 <=> v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) [avatar definition]
66071. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) <- (~33) [avatar component clause 66070] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66072. v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),sK16)) <- (33) [avatar component clause 66070] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66073. 33 | ~11 [avatar split clause 72,65700,66070]
66074. 11 | ~33 [avatar split clause 71,66070,65700]
66076. 34 <=> v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) [avatar definition]
66077. ~v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) <- (~34) [avatar component clause 66076] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66078. v1_xboole_0(k8_gate_1(sK18,k1_gate_1(sK17),k1_gate_1(sK16))) <- (34) [avatar component clause 66076] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66079. 34 | ~9 [avatar split clause 70,65692,66076]
66080. 9 | ~34 [avatar split clause 69,66076,65692]
66082. 35 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) [avatar definition]
66083. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) <- (~35) [avatar component clause 66082] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66084. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,sK16)) <- (35) [avatar component clause 66082] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66085. 35 | ~7 [avatar split clause 68,65684,66082]
66086. 7 | ~35 [avatar split clause 67,66082,65684]
66088. 36 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) [avatar definition]
66089. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) <- (~36) [avatar component clause 66088] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66090. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),sK17,k1_gate_1(sK16))) <- (36) [avatar component clause 66088] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66091. 36 | ~5 [avatar split clause 66,65676,66088]
66092. 5 | ~36 [avatar split clause 65,66088,65676]
66094. 37 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) [avatar definition]
66095. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) <- (~37) [avatar component clause 66094] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66096. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),sK16)) <- (37) [avatar component clause 66094] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66097. 37 | ~3 [avatar split clause 64,65668,66094]
66098. 3 | ~37 [avatar split clause 63,66094,65668]
66100. 38 <=> v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) [avatar definition]
66101. ~v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) <- (~38) [avatar component clause 66100] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66102. v1_xboole_0(k8_gate_1(k1_gate_1(sK18),k1_gate_1(sK17),k1_gate_1(sK16))) <- (38) [avatar component clause 66100] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66103. 38 | ~1 [avatar split clause 62,65660,66100]
66104. 1 | ~38 [avatar split clause 61,66100,65660]
66108. 39 <=> v1_xboole_0(sK16) [avatar definition]
66109. v1_xboole_0(sK16) <- (39) [avatar component clause 66108] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66110. ~v1_xboole_0(sK16) <- (~39) [avatar component clause 66108] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66113. ~v1_xboole_0(k3_gate_1(X1,X0)) | v1_xboole_0(X1) [superposition 65648,65645] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
66117. v1_xboole_0(k2_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 65649,65646] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
66119. ~v1_xboole_0(sK21) <- (~23) [resolution 65655,66011] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66120. ~v1_xboole_0(sK21) <- (~24) [resolution 65655,66017] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66121. ~v1_xboole_0(k1_gate_1(sK21)) <- (~27) [resolution 65655,66035] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66122. ~v1_xboole_0(k1_gate_1(sK21)) <- (~28) [resolution 65655,66041] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66123. ~v1_xboole_0(sK18) <- (~34) [resolution 65655,66077] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66124. ~v1_xboole_0(k1_gate_1(sK18)) <- (~38) [resolution 65655,66101] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66125. ~v1_xboole_0(k1_gate_1(sK18)) <- (~37) [resolution 65655,66095] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66127. 40 <=> v1_xboole_0(k1_gate_1(sK21)) [avatar definition]
66128. v1_xboole_0(k1_gate_1(sK21)) <- (40) [avatar component clause 66127] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66129. ~v1_xboole_0(k1_gate_1(sK21)) <- (~40) [avatar component clause 66127] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66130. ~40 | 27 [avatar split clause 66121,66034,66127]
66131. ~40 | 28 [avatar split clause 66122,66040,66127]
66133. 41 <=> v1_xboole_0(sK18) [avatar definition]
66134. v1_xboole_0(sK18) <- (41) [avatar component clause 66133] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66135. ~v1_xboole_0(sK18) <- (~41) [avatar component clause 66133] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66136. ~41 | 34 [avatar split clause 66123,66076,66133]
66138. 42 <=> v1_xboole_0(k1_gate_1(sK18)) [avatar definition]
66139. v1_xboole_0(k1_gate_1(sK18)) <- (42) [avatar component clause 66138] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66140. ~v1_xboole_0(k1_gate_1(sK18)) <- (~42) [avatar component clause 66138] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66141. ~42 | 38 [avatar split clause 66124,66100,66138]
66142. ~42 | 37 [avatar split clause 66125,66094,66138]
66143. ~v1_xboole_0(sK20) <- (~23) [resolution 65656,66011] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66144. ~v1_xboole_0(sK20) <- (~24) [resolution 65656,66017] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66147. ~v1_xboole_0(k1_gate_1(sK17)) <- (~34) [resolution 65656,66077] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66148. ~v1_xboole_0(k1_gate_1(sK17)) <- (~38) [resolution 65656,66101] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66149. ~v1_xboole_0(k1_gate_1(sK17)) <- (~37) [resolution 65656,66095] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66151. 43 <=> v1_xboole_0(k1_gate_1(sK17)) [avatar definition]
66152. v1_xboole_0(k1_gate_1(sK17)) <- (43) [avatar component clause 66151] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66153. ~v1_xboole_0(k1_gate_1(sK17)) <- (~43) [avatar component clause 66151] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66154. ~43 | 34 [avatar split clause 66147,66076,66151]
66155. ~43 | 38 [avatar split clause 66148,66100,66151]
66156. ~43 | 37 [avatar split clause 66149,66094,66151]
66158. v1_xboole_0(sK21) <- (~40) [resolution 66129,65635] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66162. 18 | 40 [avatar split clause 66158,66127,65984]
66163. ~18 | 24 [avatar split clause 66120,66016,65984]
66164. ~18 | 23 [avatar split clause 66119,66010,65984]
66165. ~v1_xboole_0(sK21) <- (40) [resolution 66128,65636] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66167. v1_xboole_0(sK18) <- (~42) [resolution 66140,65635] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66171. 41 | 42 [avatar split clause 66167,66138,66133]
66174. ~v1_xboole_0(sK19) <- (~23) [resolution 65657,66011] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66176. ~v1_xboole_0(k1_gate_1(sK16)) <- (~38) [resolution 65657,66101] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66177. ~v1_xboole_0(sK16) <- (~37) [resolution 65657,66095] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66184. ~21 | 38 [avatar split clause 66176,66100,66000]
66185. ~22 | 23 [avatar split clause 66174,66010,66004]
66187. 44 <=> v1_xboole_0(k1_gate_1(sK19)) [avatar definition]
66188. v1_xboole_0(k1_gate_1(sK19)) <- (44) [avatar component clause 66187] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66189. ~v1_xboole_0(k1_gate_1(sK19)) <- (~44) [avatar component clause 66187] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66197. ~39 | 37 [avatar split clause 66177,66094,66108]
66202. 45 <=> v1_xboole_0(sK17) [avatar definition]
66203. ~v1_xboole_0(sK17) <- (~45) [avatar component clause 66202] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66204. v1_xboole_0(sK17) <- (45) [avatar component clause 66202] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66224. v1_xboole_0(sK17) <- (~43) [resolution 66153,65635] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66228. ~20 | 24 [avatar split clause 66144,66016,65994]
66229. ~20 | 23 [avatar split clause 66143,66010,65994]
66230. 45 | 43 [avatar split clause 66224,66151,66202]
66238. ~v1_xboole_0(sK17) <- (43) [resolution 66152,65636] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66240. v1_xboole_0(sK19) <- (~44) [resolution 66189,65635] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66244. 22 | 44 [avatar split clause 66240,66187,66004]
66247. ~45 | ~43 [avatar split clause 66238,66151,66202]
66248. ~v1_xboole_0(sK16) <- (21) [resolution 66002,65636] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66249. ~39 | ~21 [avatar split clause 66248,66000,66108]
66256. 46 <=> v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) [avatar definition]
66257. v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) <- (46) [avatar component clause 66256] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66258. ~v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) <- (~46) [avatar component clause 66256] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66260. 47 <=> v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) [avatar definition]
66261. v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) <- (47) [avatar component clause 66260] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66262. ~v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) <- (~47) [avatar component clause 66260] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66271. v1_xboole_0(sK16) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (37) [resolution 65658,66096] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66272. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (38) [resolution 65658,66102] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66273. v1_xboole_0(sK19) | v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (27) [resolution 65658,66036] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66275. v1_xboole_0(sK16) | v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (31) [resolution 65658,66060] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66278. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (34) [resolution 65658,66078] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66280. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(sK21) <- (26) [resolution 65658,66030] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66281. v1_xboole_0(sK19) | v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (23) [resolution 65658,66012] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66282. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (37, ~39) [subsumption resolution 66271,66110] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66288. 48 <=> v1_xboole_0(k1_gate_1(sK20)) [avatar definition]
66289. ~v1_xboole_0(k1_gate_1(sK20)) <- (~48) [avatar component clause 66288] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66290. v1_xboole_0(k1_gate_1(sK20)) <- (48) [avatar component clause 66288] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66293. ~18 | ~40 [avatar split clause 66165,66127,65984]
66294. 42 | 43 | ~37 | 39 [avatar split clause 66282,66108,66094,66151,66138]
66301. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK20) | v1_xboole_0(sK21) <- (24) [resolution 66018,65658] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66303. ~v1_xboole_0(sK18) <- (42) [resolution 66139,65636] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66307. ~41 | ~42 [avatar split clause 66303,66138,66133]
66312. ~v1_xboole_0(sK19) <- (~25) [resolution 66023,65657] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66314. ~v1_xboole_0(sK21) <- (~25) [resolution 66023,65655] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66317. $false <- (22, ~25) [subsumption resolution 66312,66005] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66318. ~22 | 25 [avatar contradiction clause 66317]
66319. $false <- (18, ~25) [subsumption resolution 66314,65985] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66320. ~18 | 25 [avatar contradiction clause 66319]
66321. v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(sK21) <- (25) [resolution 66024,65658] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66323. ~v1_xboole_0(k1_gate_1(sK17)) <- (~33) [resolution 66071,65656] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66325. ~43 | 33 [avatar split clause 66323,66070,66151]
66327. ~v1_xboole_0(sK20) <- (~28) [resolution 66041,65656] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66330. ~v1_xboole_0(k1_gate_1(sK19)) <- (~30) [resolution 66053,65657] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66331. ~v1_xboole_0(k1_gate_1(sK20)) <- (~30) [resolution 66053,65656] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66332. ~v1_xboole_0(k1_gate_1(sK21)) <- (~30) [resolution 66053,65655] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66333. ~48 | 30 [avatar split clause 66331,66052,66288]
66335. v1_xboole_0(sK20) <- (~48) [resolution 66289,65635] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66339. 20 | 48 [avatar split clause 66335,66288,65994]
66340. ~20 | 28 [avatar split clause 66327,66040,65994]
66343. ~v1_xboole_0(sK17) | v1_xboole_0(sK16) <- (19) [resolution 65992,65642] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66348. ~v1_xboole_0(sK17) <- (19, ~39) [subsumption resolution 66343,66110] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66349. ~45 | ~19 | 39 [avatar split clause 66348,66108,65990,66202]
66350. v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (31, ~39) [subsumption resolution 66275,66110] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66353. 41 | 45 | ~31 | 39 [avatar split clause 66350,66108,66058,66202,66133]
66356. v1_xboole_0(sK16) | v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (33) [resolution 66072,65658] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66357. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (33, ~39) [subsumption resolution 66356,66110] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66358. 41 | 43 | ~33 | 39 [avatar split clause 66357,66108,66070,66151,66133]
66360. ~v1_xboole_0(sK17) <- (~35) [resolution 66083,65656] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66361. ~v1_xboole_0(k1_gate_1(sK18)) <- (~35) [resolution 66083,65655] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66362. ~45 | 35 [avatar split clause 66360,66082,66202]
66363. ~42 | 35 [avatar split clause 66361,66082,66138]
66375. 49 <=> v1_xboole_0(k4_gate_1(sK17,sK18)) [avatar definition]
66376. ~v1_xboole_0(k4_gate_1(sK17,sK18)) <- (~49) [avatar component clause 66375] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66377. v1_xboole_0(k4_gate_1(sK17,sK18)) <- (49) [avatar component clause 66375] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66384. ~v1_xboole_0(sK18) | v1_xboole_0(sK17) <- (49) [resolution 66377,65642] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66386. v1_xboole_0(sK17) <- (41, 49) [subsumption resolution 66384,66134] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66387. $false <- (41, ~45, 49) [subsumption resolution 66386,66203] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66388. ~41 | 45 | ~49 [avatar contradiction clause 66387]
66389. ~40 | 30 [avatar split clause 66332,66052,66127]
66390. ~44 | 30 [avatar split clause 66330,66052,66187]
66393. 18 | 48 | 44 | ~26 [avatar split clause 66280,66028,66187,66288,65984]
66394. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (30) [resolution 66054,65658] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66395. 40 | 48 | 44 | ~30 [avatar split clause 66394,66052,66187,66288,66127]
66401. ~v1_xboole_0(sK20) <- (48) [resolution 66290,65636] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66405. ~20 | ~48 [avatar split clause 66401,66288,65994]
66408. v1_xboole_0(sK17) | ~v1_xboole_0(sK16) <- (19) [resolution 65992,65643] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66414. ~v1_xboole_0(k1_gate_1(sK19)) <- (~26) [resolution 66029,65657] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66415. ~v1_xboole_0(k1_gate_1(sK20)) <- (~26) [resolution 66029,65656] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66416. ~v1_xboole_0(sK21) <- (~26) [resolution 66029,65655] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66417. ~44 | 26 [avatar split clause 66414,66028,66187]
66418. ~48 | 26 [avatar split clause 66415,66028,66288]
66422. ~v1_xboole_0(sK19) <- (~29) [resolution 66047,65657] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66424. ~v1_xboole_0(k1_gate_1(sK21)) <- (~29) [resolution 66047,65655] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66427. $false <- (22, ~29) [subsumption resolution 66422,66005] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66428. ~22 | 29 [avatar contradiction clause 66427]
66429. $false <- (~29, 40) [subsumption resolution 66424,66128] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66430. 29 | ~40 [avatar contradiction clause 66429]
66434. ~v1_xboole_0(sK18) <- (~31) [resolution 66059,65655] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66436. $false <- (~31, 41) [subsumption resolution 66434,66134] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66437. 31 | ~41 [avatar contradiction clause 66436]
66442. ~v1_xboole_0(k1_gate_1(sK16)) <- (~32) [resolution 66065,65657] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66444. ~v1_xboole_0(sK18) <- (~32) [resolution 66065,65655] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66447. $false <- (21, ~32) [subsumption resolution 66442,66002] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66448. ~21 | 32 [avatar contradiction clause 66447]
66449. $false <- (~32, 41) [subsumption resolution 66444,66134] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66450. 32 | ~41 [avatar contradiction clause 66449]
66452. ~v1_xboole_0(k4_gate_1(sK17,sK18)) <- (~47) [resolution 66262,66117] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66454. ~49 | 47 [avatar split clause 66452,66260,66375]
66455. v1_xboole_0(sK19) | v1_xboole_0(k1_gate_1(sK20)) | v1_xboole_0(k1_gate_1(sK21)) <- (29) [resolution 66048,65658] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66457. v1_xboole_0(sK18) | v1_xboole_0(sK17) <- (~49) [resolution 66376,65639] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66459. ~v1_xboole_0(k1_gate_1(sK16)) <- (~36) [resolution 66089,65657] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66463. $false <- (21, ~36) [subsumption resolution 66459,66002] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66464. ~21 | 36 [avatar contradiction clause 66463]
66465. v1_xboole_0(sK18) <- (~45, ~49) [subsumption resolution 66457,66203] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66466. 41 | 45 | 49 [avatar split clause 66465,66375,66202,66133]
66468. ~v1_xboole_0(sK16) <- (~31) [resolution 66059,65657] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66481. ~v1_xboole_0(k1_gate_1(sK16)) <- (~46) [resolution 66258,66117] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66484. $false <- (21, ~46) [subsumption resolution 66481,66002] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66485. ~21 | 46 [avatar contradiction clause 66484]
66486. ~18 | 26 [avatar split clause 66416,66028,65984]
66491. ~39 | 31 [avatar split clause 66468,66058,66108]
66492. ~v1_xboole_0(sK16) <- (19, ~45) [subsumption resolution 66408,66203] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66493. ~39 | ~19 | 45 [avatar split clause 66492,66202,65990,66108]
66494. 40 | 20 | 22 | ~27 [avatar split clause 66273,66034,66004,65994,66127]
66495. 18 | 20 | 22 | ~23 [avatar split clause 66281,66010,66004,65994,65984]
66497. 18 | 20 | 44 | ~24 [avatar split clause 66301,66016,66187,65994,65984]
66498. 18 | 48 | 22 | ~25 [avatar split clause 66321,66022,66004,66288,65984]
66499. 40 | 48 | 22 | ~29 [avatar split clause 66455,66046,66004,66288,66127]
66500. v1_xboole_0(sK16) <- (~21) [resolution 66001,65635] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66501. 39 | 21 [avatar split clause 66500,66000,66108]
66518. ~v1_xboole_0(k1_gate_1(sK18)) <- (~36) [resolution 66089,65655] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66522. ~42 | 36 [avatar split clause 66518,66088,66138]
66523. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (32) [resolution 66066,65658] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66524. v1_xboole_0(sK17) | v1_xboole_0(sK18) <- (~21, 32) [subsumption resolution 66523,66001] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66525. v1_xboole_0(sK18) <- (~21, 32, ~45) [subsumption resolution 66524,66203] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66526. 41 | 21 | ~32 | 45 [avatar split clause 66525,66202,66064,66000,66133]
66528. ~v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) | ~v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) <- (~17) [resolution 65981,65653] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66529. ~46 | ~47 | 17 [avatar split clause 66528,65980,66260,66256]
66533. ~v1_xboole_0(sK18) <- (~46) [resolution 66258,65649] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66535. $false <- (41, ~46) [subsumption resolution 66533,66134] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66536. ~41 | 46 [avatar contradiction clause 66535]
66541. ~v1_xboole_0(sK16) <- (~47) [resolution 66262,65649] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66543. $false <- (39, ~47) [subsumption resolution 66541,66109] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66544. ~39 | 47 [avatar contradiction clause 66543]
66547. ~v1_xboole_0(k1_gate_1(sK20)) <- (~25) [resolution 66023,65656] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66552. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(k1_gate_1(sK18)) <- (~21, 38) [subsumption resolution 66272,66001] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66553. 42 | 43 | 21 | ~38 [avatar split clause 66552,66100,66000,66151,66138]
66554. v1_xboole_0(k1_gate_1(sK17)) | v1_xboole_0(sK18) <- (~21, 34) [subsumption resolution 66278,66001] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66555. v1_xboole_0(k1_gate_1(sK17)) <- (~21, 34, ~41) [subsumption resolution 66554,66135] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66556. 43 | 21 | ~34 | 41 [avatar split clause 66555,66133,66076,66000,66151]
66557. ~48 | 25 [avatar split clause 66547,66022,66288]
66564. ~v1_xboole_0(sK17) <- (~32) [resolution 66065,65656] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66567. $false <- (~32, 45) [subsumption resolution 66564,66204] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66568. 32 | ~45 [avatar contradiction clause 66567]
66573. ~v1_xboole_0(k1_gate_1(sK20)) <- (~29) [resolution 66047,65656] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66575. ~48 | 29 [avatar split clause 66573,66046,66288]
66582. v1_xboole_0(k2_gate_1(sK18,k1_gate_1(sK16))) <- (17) [resolution 65982,66113] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66583. v1_xboole_0(k2_gate_1(sK16,k4_gate_1(sK17,sK18))) <- (17) [resolution 65982,65648] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66584. 46 | ~17 [avatar split clause 66582,65980,66256]
66585. 47 | ~17 [avatar split clause 66583,65980,66260]
66587. v1_xboole_0(sK18) | ~v1_xboole_0(sK17) <- (49) [resolution 66377,65643] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66590. ~v1_xboole_0(sK17) <- (~41, 49) [subsumption resolution 66587,66135] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66591. $false <- (~41, 45, 49) [subsumption resolution 66590,66204] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66592. 41 | ~45 | ~49 [avatar contradiction clause 66591]
66594. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK18) <- (46) [resolution 66257,65650] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66596. v1_xboole_0(sK18) <- (~21, 46) [subsumption resolution 66594,66001] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66599. 41 | 21 | ~46 [avatar split clause 66596,66256,66000,66133]
66602. ~v1_xboole_0(sK20) <- (~27) [resolution 66035,65656] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66605. $false <- (20, ~27) [subsumption resolution 66602,65995] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66606. ~20 | 27 [avatar contradiction clause 66605]
66610. v1_xboole_0(k4_gate_1(sK17,sK18)) | v1_xboole_0(sK16) <- (47) [resolution 66261,65650] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66613. ~v1_xboole_0(sK17) <- (~36) [resolution 66089,65656] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66616. $false <- (~36, 45) [subsumption resolution 66613,66204] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66617. 36 | ~45 [avatar contradiction clause 66616]
66620. ~v1_xboole_0(sK16) <- (~33) [resolution 66071,65657] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66622. ~v1_xboole_0(sK18) <- (~33) [resolution 66071,65655] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66625. $false <- (~33, 39) [subsumption resolution 66620,66109] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66626. 33 | ~39 [avatar contradiction clause 66625]
66627. $false <- (~33, 41) [subsumption resolution 66622,66134] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66628. 33 | ~41 [avatar contradiction clause 66627]
66629. ~v1_xboole_0(sK16) | ~v1_xboole_0(sK17) <- (~19) [resolution 65991,65640] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66630. v1_xboole_0(sK17) | v1_xboole_0(sK16) <- (~19) [resolution 65991,65639] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66631. ~v1_xboole_0(sK19) <- (~27) [resolution 66035,65657] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66639. v1_xboole_0(k1_gate_1(sK16)) | v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (36) [resolution 66090,65658] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66641. v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (~21, 36) [subsumption resolution 66639,66001] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66642. v1_xboole_0(k1_gate_1(sK18)) <- (~21, 36, ~45) [subsumption resolution 66641,66203] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66643. $false <- (~21, 36, ~42, ~45) [subsumption resolution 66642,66140] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66644. 21 | ~36 | 42 | 45 [avatar contradiction clause 66643]
66646. ~v1_xboole_0(sK16) <- (~35) [resolution 66083,65657] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66650. $false <- (~35, 39) [subsumption resolution 66646,66109] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66651. 35 | ~39 [avatar contradiction clause 66650]
66652. v1_xboole_0(sK16) | v1_xboole_0(sK17) | v1_xboole_0(k1_gate_1(sK18)) <- (35) [resolution 66084,65658] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66654. ~v1_xboole_0(k1_gate_1(sK19)) <- (~28) [resolution 66041,65657] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66660. ~22 | 27 [avatar split clause 66631,66034,66004]
66661. v1_xboole_0(sK16) | v1_xboole_0(sK17) <- (35, ~42) [subsumption resolution 66652,66140] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66662. 45 | 39 | ~35 | 42 [avatar split clause 66661,66138,66082,66108,66202]
66663. 39 | 45 | 19 [avatar split clause 66630,65990,66202,66108]
66664. 39 | 49 | ~47 [avatar split clause 66610,66260,66375,66108]
66665. ~44 | 28 [avatar split clause 66654,66040,66187]
66666. ~45 | ~39 | 19 [avatar split clause 66629,65990,66108,66202]
66669. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(sK20) | v1_xboole_0(k1_gate_1(sK21)) <- (28) [resolution 66042,65658] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66670. v1_xboole_0(k1_gate_1(sK19)) | v1_xboole_0(k1_gate_1(sK21)) <- (~20, 28) [subsumption resolution 66669,65996] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66679. ~v1_xboole_0(sK19) <- (44) [resolution 66188,65636] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66681. $false <- (22, 44) [subsumption resolution 66679,66005] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
66682. ~22 | ~44 [avatar contradiction clause 66681]
66683. 40 | 44 | 20 | ~28 [avatar split clause 66670,66040,65994,66187,66127]
66684. ~v1_xboole_0(k1_gate_1(sK19)) <- (~24) [resolution 66017,65657] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66687. ~44 | 24 [avatar split clause 66684,66016,66187]
66692. ~v1_xboole_0(k1_gate_1(sK16)) <- (~34) [resolution 66077,65657] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66696. $false <- (21, ~34) [subsumption resolution 66692,66002] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66697. ~21 | 34 [avatar contradiction clause 66696]
66701. ~v1_xboole_0(sK17) <- (~31) [resolution 66059,65656] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66704. $false <- (~31, 45) [subsumption resolution 66701,66204] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
66705. 31 | ~45 [avatar contradiction clause 66704]
66716. ~v1_xboole_0(sK17) | ~v1_xboole_0(sK18) <- (~49) [resolution 66376,65640] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66719. ~v1_xboole_0(sK18) <- (45, ~49) [subsumption resolution 66716,66204] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66720. $false <- (41, 45, ~49) [subsumption resolution 66719,66134] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
66721. ~41 | ~45 | 49 [avatar contradiction clause 66720]
66722. $false [avatar sat refutation 65850,65914,65946,65962,65970,65974,65976,65977,65987,65988,65997,65998,66007,66008,66013,66014,66019,66020,66025,66026,66031,66032,66037,66038,66043,66044,66049,66050,66055,66056,66061,66062,66067,66068,66073,66074,66079,66080,66085,66086,66091,66092,66097,66098,66103,66104,66130,66131,66136,66141,66142,66154,66155,66156,66162,66163,66164,66171,66184,66185,66197,66228,66229,66230,66244,66247,66249,66293,66294,66307,66318,66320,66325,66333,66339,66340,66349,66353,66358,66362,66363,66388,66389,66390,66393,66395,66405,66417,66418,66428,66430,66437,66448,66450,66454,66464,66466,66485,66486,66491,66493,66494,66495,66497,66498,66499,66501,66522,66526,66529,66536,66544,66553,66556,66557,66568,66575,66584,66585,66592,66599,66606,66617,66626,66628,66644,66651,66660,66662,66663,66664,66665,66666,66682,66683,66687,66697,66705,66721] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:5738,thDist:-5738}
% SZS output end Proof for gate_2__t1_gate_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 28
% >>> Preprocessing
% Introduced skolems: 22
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 65598
% Generated clauses: 66416
% Active clauses: 302
% Passive clauses: 361
% Final active clauses: 68
% Final passive clauses: 5
% Main loop iterations started: 307
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 79
% 
% >>> Deletion Inferences
% Simple tautologies: 65294
% Forward subsumptions: 210
% 
% >>> Generating Inferences
% Binary resolution: 254
% Unit resulting resolution: 39
% Forward superposition: 16
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 325
% Split components: 4242
% Sat splitting refutations: 78
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 472
% SAT solver binary clauses: 155
% 
% Memory used [KB]: 31598
% Time elapsed: 0.116 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% property evaluation: 0.046 s
% preprocessing: 0.010 s
% term sharing: 0.001 s
% splitting model update: 0.001 s ( own 0.0000 s ) 
% unit clause index maintenance: 0.001 s
% binary resolution index maintenance: 0.001 s
% forward demodulation: 0.001 s
% unit resulting resolution: 0.004 s ( own 0.003 s ) 
% other: 0.116 s ( own 0.052 s ) 

