i: [1,36,0,3,1,1,0,0] t5_zfmisc_1
a: [3,51,0,3,1,36]
s: 51
i: [1,35,0,6,1,1,0,0] t5_zfmisc_1
a: [3,56,0,6,1,35]
s: 56
i: [1,37,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,61,0,2,1,37]
s: 61
i: [1,38,0,2,1,0,0,255] fc1_xboole_0
a: [3,66,0,2,1,38]
s: 66
i: [1,46,0,2,1,0,0,2] rc2_xboole_0
a: [3,76,0,2,1,46]
s: 76
i: [1,47,0,2,1,0,0,2] rc1_xboole_0
a: [3,81,0,2,1,47]
s: 81
i: [1,40,0,3,1,0,0,1] fc2_xboole_0
s: 40
i: [1,39,0,3,1,0,0,255] d2_xboole_0
a: [3,71,0,3,1,39]
s: 71
r: 61
i: [1,44,0,4,1,0,0,1] fc3_xboole_0
s: 44
i: [1,41,0,5,2,0,0,255] t6_boole
s: 41
i: [1,42,0,7,3,0,0,1] t8_boole
s: 42
i: [1,45,0,7,1,0,0,1] commutativity_k2_tarski
s: 45
i: [1,43,0,9,2,0,0,1] t4_zfmisc_1
s: 43
d: [2,89,1,8,2,1,44,43,56]
s: 89
d: [2,85,1,7,3,0,44,41,41]
s: 85
d: [2,90,1,9,2,0,44,43,45]
s: 90
d: [2,92,2,3,1,1,47,89]
a: [3,96,2,3,1,92]
s: 96
r: 89
r: 56
r: 51
d: [2,103,0,3,1,2,23,51,96]
a: [3,112,0,3,1,103]
s: 112
d: [2,102,0,6,1,2,23,56,96]
a: [3,107,0,6,1,102]
s: 107
d: [2,114,2,0,0,2,52,112,107,90]
e: 114
f: 114,103,102,92,35,36
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t5_zfmisc_1
% SZS output start Proof for zfmisc_1__t5_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (k1_tarski(X0) = k2_tarski(X1,X2) => X1 = X2) [input t5_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (k1_tarski(X0) = k2_tarski(X1,X2) => X1 = X2) [negated conjecture 1]
4. ! [X0,X1] : k2_tarski(X0,X1) = k2_tarski(X1,X0) [input commutativity_k2_tarski]
15. ! [X0] : ! [X1] : ! [X2] : (k1_tarski(X0) = k2_tarski(X1,X2) => X0 = X1) [input t4_zfmisc_1]
19. ~! [X0,X1,X2] : (k1_tarski(X0) = k2_tarski(X1,X2) => X1 = X2) [flattening 2]
22. ! [X0,X1,X2] : (k1_tarski(X0) = k2_tarski(X1,X2) => X0 = X1) [flattening 15]
25. ? [X0,X1,X2] : (X1 != X2 & k1_tarski(X0) = k2_tarski(X1,X2)) [ennf transformation 19]
28. ! [X0,X1,X2] : (X0 = X1 | k1_tarski(X0) != k2_tarski(X1,X2)) [ennf transformation 22]
29. ? [X0,X1,X2] : (X1 != X2 & k1_tarski(X0) = k2_tarski(X1,X2)) => (sK1 != sK2 & k1_tarski(sK0) = k2_tarski(sK1,sK2)) [choice axiom]
30. sK1 != sK2 & k1_tarski(sK0) = k2_tarski(sK1,sK2) [skolemisation 25,29]
35. k1_tarski(sK0) = k2_tarski(sK1,sK2) [cnf transformation 30] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. sK1 != sK2 [cnf transformation 30] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. k1_tarski(X0) != k2_tarski(X1,X2) | X0 = X1 [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. k2_tarski(X0,X1) = k2_tarski(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. 1 <=> sK1 = sK2 [avatar definition]
51. sK1 != sK2 <- (~1) [avatar component clause 49] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
52. ~1 [avatar split clause 36,49]
54. 2 <=> k1_tarski(sK0) = k2_tarski(sK1,sK2) [avatar definition]
56. k1_tarski(sK0) = k2_tarski(sK1,sK2) <- (2) [avatar component clause 54] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
57. 2 [avatar split clause 35,54]
89. k1_tarski(X0) != k1_tarski(sK0) | sK1 = X0 <- (2) [superposition 43,56] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
90. k2_tarski(X2,X1) != k1_tarski(X3) | X1 = X3 [superposition 43,45] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
92. sK0 = sK1 <- (2) [equality resolution 89] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
94. 8 <=> sK0 = sK1 [avatar definition]
96. sK0 = sK1 <- (8) [avatar component clause 94] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
97. 8 | ~2 [avatar split clause 92,54,94]
102. k1_tarski(sK0) = k2_tarski(sK0,sK2) <- (2, 8) [backward demodulation 56,96] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
103. sK0 != sK2 <- (~1, 8) [backward demodulation 51,96] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
105. 9 <=> k1_tarski(sK0) = k2_tarski(sK0,sK2) [avatar definition]
107. k1_tarski(sK0) = k2_tarski(sK0,sK2) <- (9) [avatar component clause 105] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
108. 9 | ~2 | ~8 [avatar split clause 102,94,54,105]
110. 10 <=> sK0 = sK2 [avatar definition]
112. sK0 != sK2 <- (~10) [avatar component clause 110] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
113. ~10 | 1 | ~8 [avatar split clause 103,94,49,110]
114. $false <- (9, ~10) [unit resulting resolution 112,107,90] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
118. ~9 | 10 [avatar contradiction clause 114]
119. $false [avatar sat refutation 52,57,97,108,113,118] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:39,thDist:-39}
% SZS output end Proof for zfmisc_1__t5_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 43
% Active clauses: 18
% Passive clauses: 22
% Final active clauses: 14
% Final passive clauses: 3
% Main loop iterations started: 19
% 
% >>> Simplifying Inferences
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Unit resulting resolution: 1
% Forward superposition: 7
% Backward superposition: 6
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 7
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

