i: [1,233,0,2,1,1,0,0] t33_xreal_1
a: [3,365,0,2,1,233]
s: 365
i: [1,232,0,2,1,1,0,0] t33_xreal_1
a: [3,370,0,2,1,232]
s: 370
i: [1,235,0,3,1,1,0,0] t33_xreal_1
a: [3,355,0,3,1,235]
s: 355
i: [1,234,0,3,1,1,0,0] t33_xreal_1
a: [3,360,0,3,1,234]
s: 360
i: [1,236,0,5,1,1,0,0] t33_xreal_1
a: [3,350,0,5,1,236]
s: 350
i: [1,237,0,2,1,0,0,255] fc1_numbers
a: [3,375,0,2,1,237]
s: 375
i: [1,238,0,2,1,0,0,2] fc1_xboole_0
a: [3,380,0,2,1,238]
s: 380
i: [1,245,0,2,1,0,0,3] fc6_ordinal1
a: [3,410,0,2,1,245]
s: 410
i: [1,244,0,2,1,0,0,3] fc6_ordinal1
a: [3,415,0,2,1,244]
s: 415
i: [1,314,0,2,1,0,0,2] rc2_xboole_0
a: [3,427,0,2,1,314]
s: 427
i: [1,315,0,2,1,0,0,4] rc4_ordinal1
a: [3,432,0,2,1,315]
s: 432
i: [1,316,0,2,1,0,0,4] rc1_ordinal1
a: [3,437,0,2,1,316]
s: 437
i: [1,317,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,442,0,2,1,317]
s: 442
i: [1,318,0,2,1,0,0,1] rc1_xreal_0
a: [3,447,0,2,1,318]
s: 447
i: [1,319,0,2,1,0,0,2] rc1_xxreal_0
a: [3,452,0,2,1,319]
s: 452
i: [1,320,0,2,1,0,0,2] rc1_xboole_0
a: [3,457,0,2,1,320]
s: 457
i: [1,324,0,2,1,0,0,255] rc2_ordinal1
a: [3,462,0,2,1,324]
s: 462
i: [1,323,0,2,1,0,0,255] rc2_ordinal1
a: [3,467,0,2,1,323]
s: 467
i: [1,322,0,2,1,0,0,255] rc2_ordinal1
a: [3,472,0,2,1,322]
s: 472
i: [1,321,0,2,1,0,0,255] rc2_ordinal1
a: [3,477,0,2,1,321]
s: 477
i: [1,326,0,2,1,0,0,4] rc5_ordinal1
a: [3,482,0,2,1,326]
s: 482
i: [1,325,0,2,1,0,0,4] rc5_ordinal1
a: [3,487,0,2,1,325]
s: 487
i: [1,328,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,492,0,2,1,328]
s: 492
i: [1,327,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,497,0,2,1,327]
s: 497
i: [1,332,0,2,1,0,0,2] rc3_xreal_0
a: [3,502,0,2,1,332]
s: 502
i: [1,331,0,2,1,0,0,2] rc3_xreal_0
a: [3,507,0,2,1,331]
s: 507
i: [1,330,0,2,1,0,0,2] rc3_xreal_0
a: [3,512,0,2,1,330]
s: 512
i: [1,329,0,2,1,0,0,2] rc3_xreal_0
a: [3,517,0,2,1,329]
s: 517
i: [1,336,0,2,1,0,0,2] rc2_xreal_0
a: [3,522,0,2,1,336]
s: 522
i: [1,335,0,2,1,0,0,2] rc2_xreal_0
a: [3,527,0,2,1,335]
s: 527
i: [1,334,0,2,1,0,0,2] rc2_xreal_0
a: [3,532,0,2,1,334]
s: 532
i: [1,333,0,2,1,0,0,2] rc2_xreal_0
a: [3,537,0,2,1,333]
s: 537
i: [1,338,0,2,1,0,0,2] rc2_xxreal_0
a: [3,542,0,2,1,338]
s: 542
i: [1,337,0,2,1,0,0,2] rc2_xxreal_0
a: [3,547,0,2,1,337]
s: 547
i: [1,340,0,2,1,0,0,2] rc3_xxreal_0
a: [3,552,0,2,1,340]
s: 552
i: [1,339,0,2,1,0,0,2] rc3_xxreal_0
a: [3,557,0,2,1,339]
s: 557
i: [1,344,0,2,1,0,0,2] rc4_xreal_0
a: [3,562,0,2,1,344]
s: 562
i: [1,343,0,2,1,0,0,2] rc4_xreal_0
a: [3,567,0,2,1,343]
s: 567
i: [1,342,0,2,1,0,0,2] rc4_xreal_0
a: [3,572,0,2,1,342]
s: 572
i: [1,341,0,2,1,0,0,2] rc4_xreal_0
a: [3,577,0,2,1,341]
s: 577
i: [1,346,0,2,1,0,0,2] rc4_xxreal_0
a: [3,582,0,2,1,346]
s: 582
i: [1,345,0,2,1,0,0,2] rc4_xxreal_0
a: [3,587,0,2,1,345]
s: 587
i: [1,296,0,3,1,0,0,255] reflexivity_r1_tarski
s: 296
i: [1,239,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,385,0,3,1,239]
s: 385
i: [1,240,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,390,0,3,1,240]
s: 390
r: 380
d: [2,589,0,2,1,2,23,380,390]
a: [3,593,0,2,1,589]
s: 593
i: [1,241,0,3,1,0,0,2] t1_numerals
a: [3,395,0,3,1,241]
s: 395
d: [2,595,0,3,1,2,22,395,390]
a: [3,599,0,3,1,595]
s: 599
i: [1,247,0,4,2,0,0,4] cc6_ordinal1
s: 247
i: [1,248,0,4,2,0,0,4] cc2_xcmplx_0
s: 248
i: [1,249,0,4,2,0,0,4] cc2_xreal_0
s: 249
d: [2,601,1,2,1,1,40,247,432]
a: [3,606,1,2,1,601]
s: 606
i: [1,250,0,4,2,0,0,4] cc2_xxreal_0
s: 250
d: [2,602,1,2,1,1,40,247,482]
a: [3,611,1,2,1,602]
s: 611
i: [1,251,0,4,2,0,0,255] cc1_ordinal1
s: 251
i: [1,252,0,4,2,0,0,255] cc1_ordinal1
s: 252
d: [2,613,1,2,1,1,40,248,432]
a: [3,618,1,2,1,613]
s: 618
i: [1,255,0,4,2,0,0,2] cc3_xreal_0
s: 255
d: [2,699,1,2,1,1,40,255,370]
a: [3,708,1,2,1,699]
s: 708
i: [1,256,0,4,2,0,0,2] cc4_xreal_0
s: 256
i: [1,265,0,4,2,0,0,255] cc9_ordinal1
s: 265
d: [2,700,1,2,1,1,40,255,365]
a: [3,713,1,2,1,700]
s: 713
i: [1,266,0,4,2,0,0,4] cc7_ordinal1
s: 266
d: [2,720,1,2,1,1,40,256,370]
a: [3,729,1,2,1,720]
s: 729
i: [1,267,0,4,2,0,0,4] cc3_ordinal1
s: 267
i: [1,295,0,4,1,0,0,3] existence_m1_subset_1
s: 295
d: [2,721,1,2,1,1,40,256,365]
a: [3,734,1,2,1,721]
s: 734
i: [1,242,0,4,1,0,0,255] dt_k5_numbers
a: [3,400,0,4,1,242]
s: 400
d: [2,614,1,2,1,1,40,248,482]
a: [3,623,1,2,1,614]
s: 623
i: [1,243,0,4,1,0,0,255] dt_k6_numbers
a: [3,405,0,4,1,243]
s: 405
d: [2,813,0,4,1,2,22,400,385]
a: [3,817,0,4,1,813]
s: 817
d: [2,625,1,2,1,1,40,249,432]
a: [3,630,1,2,1,625]
s: 630
d: [2,819,0,4,1,2,22,405,385]
a: [3,823,0,4,1,819]
s: 823
d: [2,626,1,2,1,1,40,249,482]
a: [3,635,1,2,1,626]
s: 635
i: [1,268,0,5,2,0,0,2] t6_boole
s: 268
i: [1,269,0,5,2,0,0,3] cc8_ordinal1
s: 269
d: [2,637,1,2,1,1,40,250,432]
a: [3,642,1,2,1,637]
s: 642
i: [1,270,0,5,2,0,0,255] cc1_xcmplx_0
s: 270
d: [2,638,1,2,1,1,40,250,482]
a: [3,647,1,2,1,638]
s: 647
i: [1,271,0,5,2,0,0,255] cc1_xreal_0
s: 271
i: [1,292,0,5,2,0,0,255] t7_boole
s: 292
d: [2,649,1,2,1,1,40,251,410]
a: [3,657,1,2,1,649]
s: 657
i: [1,307,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,422,0,5,2,307]
s: 422
d: [2,651,1,2,1,1,40,251,437]
a: [3,667,1,2,1,651]
s: 667
d: [2,831,0,5,2,1,22,268,390]
s: 831
i: [1,276,0,6,3,0,0,255] cc2_ordinal1
s: 276
d: [2,854,1,5,2,2,44,355,831]
s: 854
i: [1,279,0,6,3,0,0,2] cc3_xxreal_0
s: 279
d: [2,855,1,5,2,2,44,360,831]
s: 855
i: [1,281,0,6,3,0,0,2] cc3_xxreal_0
s: 281
i: [1,282,0,6,3,0,0,2] cc5_xxreal_0
s: 282
d: [2,853,1,7,2,2,44,350,831]
s: 853
i: [1,284,0,6,3,0,0,2] cc5_xxreal_0
s: 284
d: [2,902,2,8,2,3,40,853,422]
a: [3,906,2,4,1,902]
s: 906
i: [1,286,0,6,3,0,0,2] cc7_xxreal_0
s: 286
i: [1,287,0,6,3,0,0,2] cc7_xxreal_0
s: 287
d: [2,674,1,2,1,1,40,252,410]
a: [3,682,1,2,1,674]
s: 682
i: [1,288,0,6,2,0,0,255] t1_subset
s: 288
d: [2,676,1,2,1,1,40,252,437]
a: [3,692,1,2,1,676]
s: 692
i: [1,297,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 297
i: [1,253,0,7,3,0,0,4] cc5_ordinal1
s: 253
d: [2,701,1,2,1,1,40,255,447]
a: [3,718,1,2,1,701]
s: 718
i: [1,254,0,7,2,0,0,1] t1_arithm
s: 254
d: [2,722,1,2,1,1,40,256,447]
a: [3,739,1,2,1,722]
s: 739
i: [1,290,0,7,2,0,0,255] t3_subset
s: 290
i: [1,291,0,7,3,0,0,2] t8_boole
s: 291
d: [2,741,1,2,1,1,40,265,593]
a: [3,748,1,2,1,741]
s: 748
i: [1,246,0,8,3,0,0,255] cc10_ordinal1
s: 246
d: [2,742,1,2,1,1,40,265,457]
a: [3,753,1,2,1,742]
s: 753
i: [1,273,0,8,4,0,0,2] cc4_xxreal_0
s: 273
i: [1,275,0,8,4,0,0,2] cc6_xxreal_0
s: 275
d: [2,934,2,6,3,1,40,273,729]
a: [3,949,2,2,1,934]
s: 949
r: 949
i: [1,277,0,8,4,0,0,2] cc8_xxreal_0
s: 277
a: [3,948,2,2,1,934]
s: 948
i: [1,289,0,8,3,0,0,255] t2_subset
s: 289
i: [1,304,0,8,3,0,0,1] fc2_xcmplx_0
s: 304
a: [3,953,2,2,1,934]
s: 953
i: [1,306,0,8,3,0,0,1] fc5_xreal_0
s: 306
d: [2,935,2,6,3,1,40,273,734]
a: [3,962,2,2,1,935]
s: 962
r: 962
i: [1,294,0,9,3,0,0,255] t5_subset
s: 294
i: [1,293,0,10,3,0,0,255] t4_subset
s: 293
a: [3,961,2,2,1,935]
s: 961
i: [1,308,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 308
a: [3,966,2,2,1,935]
s: 966
i: [1,259,0,11,5,0,0,1] t2_real
s: 259
i: [1,260,0,11,5,0,0,1] t1_real
s: 260
a: [3,956,2,2,1,934]
s: 956
i: [1,261,0,11,5,0,0,1] t7_real
s: 261
a: [3,969,2,2,1,935]
s: 969
i: [1,262,0,11,5,0,0,1] t8_real
s: 262
i: [1,263,0,11,5,0,0,1] t4_real
s: 263
d: [2,1130,1,12,4,1,40,261,350]
a: [3,1139,1,4,1,1130]
s: 1139
d: [2,1153,2,0,0,3,52,370,365,1139,306]
e: 1153
d: [2,1154,2,4,2,1,40,1139,306]
d: [2,1156,2,2,1,2,21,1154,365]
d: [2,1157,2,0,0,3,21,1156,370]
e: 1157
r: 1139
i: [1,264,0,11,5,0,0,1] t3_real
s: 264
a: [3,1138,1,4,1,1130]
s: 1138
d: [2,1159,3,0,0,2,52,906,1138,256]
e: 1159
d: [2,1160,2,4,1,1,40,1138,256]
d: [2,1163,2,0,0,2,21,1160,906]
e: 1163
r: 906
i: [1,305,0,11,3,0,0,2] commutativity_k2_xcmplx_0
s: 305
i: [1,302,0,12,5,0,0,1] fc10_xreal_0
s: 302
d: [2,1088,1,11,3,1,40,308,350]
a: [3,1173,1,2,1,1088]
s: 1173
i: [1,303,0,12,5,0,0,1] fc9_xreal_0
s: 303
a: [3,1143,1,4,1,1130]
s: 1143
d: [2,1216,3,0,0,5,52,956,370,969,365,1143,303]
e: 1216
d: [2,1217,2,8,4,1,40,1143,303]
d: [2,1221,2,6,3,2,21,1217,365]
d: [2,1222,2,4,2,3,21,1221,969]
d: [2,1223,2,2,1,4,21,1222,370]
d: [2,1224,2,0,0,5,21,1223,956]
e: 1224
r: 956
i: [1,309,0,12,5,0,0,1] fc11_xreal_0
s: 309
i: [1,310,0,12,5,0,0,1] fc14_xreal_0
s: 310
d: [2,1093,1,8,4,1,40,259,355]
d: [2,1100,1,6,3,2,21,1093,365]
a: [3,1104,1,2,1,1100]
s: 1104
i: [1,311,0,12,5,0,0,1] fc13_xreal_0
s: 311
a: [3,957,2,2,1,934]
s: 957
d: [2,1267,3,0,0,3,52,729,953,957,281]
e: 1267
d: [2,1269,3,4,2,1,40,957,281]
d: [2,1271,3,2,1,2,21,1269,953]
d: [2,1272,3,0,0,3,21,1271,729]
e: 1272
r: 953
i: [1,312,0,12,5,0,0,1] fc12_xreal_0
s: 312
i: [1,257,0,13,6,0,0,1] t6_real
s: 257
a: [3,952,2,2,1,934]
s: 952
i: [1,258,0,13,6,0,0,1] t5_real
s: 258
a: [3,905,2,4,1,902]
s: 905
i: [1,301,0,14,4,0,0,255] existence_m2_subset_1
s: 301
i: [1,298,0,15,5,0,0,255] dt_m2_subset_1
s: 298
a: [3,910,2,4,1,902]
s: 910
i: [1,299,0,15,5,0,0,255] redefinition_m2_subset_1
s: 299
d: [2,1161,2,4,1,1,40,1138,255]
a: [3,1182,2,4,1,1161]
s: 1182
i: [1,300,0,15,5,0,0,255] redefinition_m2_subset_1
s: 300
i: [1,313,0,17,4,0,0,2] spc6_arithm
s: 313
d: [2,1212,2,4,2,2,44,1173,831]
s: 1212
r: 282
r: 279
d: [2,1385,3,0,0,4,52,577,567,1212]
e: 1385
d: [2,1386,3,0,0,4,52,587,582,1212]
e: 1386
d: [2,1398,3,2,1,3,40,1212,567]
d: [2,1402,3,0,0,4,21,1398,577]
e: 1402
d: [2,1399,3,2,1,3,40,1212,582]
d: [2,1404,3,0,0,4,21,1399,587]
e: 1404
r: 1212
r: 1173
s: 279
a: [3,1172,1,2,1,1088]
s: 1172
s: 282
a: [3,1177,1,5,1,1088]
s: 1177
d: [2,1219,2,8,2,1,40,1143,281]
a: [3,1230,2,4,1,1219]
s: 1230
d: [2,743,1,2,1,1,40,265,577]
a: [3,758,1,2,1,743]
s: 758
d: [2,1259,2,4,2,2,44,1104,831]
s: 1259
d: [2,1434,3,0,0,4,52,562,577,1259]
e: 1434
d: [2,1437,3,2,1,3,40,1259,577]
d: [2,1440,3,0,0,4,21,1437,562]
e: 1440
r: 1259
r: 1104
a: [3,1103,1,2,1,1100]
s: 1103
a: [3,1108,1,2,1,1100]
s: 1108
d: [2,1455,2,0,0,3,52,1172,593,1108,282]
e: 1455
d: [2,1456,2,4,2,1,40,1108,282]
d: [2,1460,2,2,1,2,21,1456,593]
d: [2,1461,2,0,0,3,21,1460,1172]
e: 1461
r: 952
r: 957
r: 1108
r: 1143
r: 1230
a: [3,1107,1,2,1,1100]
s: 1107
a: [3,1147,1,2,1,1130]
s: 1147
d: [2,1496,2,0,0,3,52,1172,593,1147,279]
e: 1496
d: [2,1497,2,4,2,1,40,1147,279]
d: [2,1500,2,2,1,2,21,1497,593]
d: [2,1501,2,0,0,3,21,1500,1172]
e: 1501
d: [2,1094,1,8,4,1,40,259,360]
d: [2,1110,1,6,3,2,21,1094,370]
f: 1501,1461,1440,1404,1223,1157,1130,1110,1100,589,345,346,341,344,240,238,232,233,234,235,236
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t33_xreal_1
% SZS output start Proof for xreal_1__t33_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0)) => r1_xxreal_0(k6_numbers,k2_xcmplx_0(X0,X1))))) [input t33_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0)) => r1_xxreal_0(k6_numbers,k2_xcmplx_0(X0,X1))))) [negated conjecture 1]
12. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
15. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
24. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k2_xcmplx_0(X0,X1))) [input fc9_xreal_0]
34. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
41. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v3_xxreal_0(X1) & ~v2_xxreal_0(X0) & ~r1_xxreal_0(X0,X1)))) [input t7_real]
59. ? [X0] : (v1_xreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0) & v1_xboole_0(X0)) [input rc4_xreal_0]
60. ? [X0] : (v1_xxreal_0(X0) & v1_xboole_0(X0)) [input rc4_xxreal_0]
74. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
81. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
85. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
90. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k2_xcmplx_0(X0,X1))) [input fc5_xreal_0]
103. ? [X0] : (? [X1] : ((~r1_xxreal_0(k6_numbers,k2_xcmplx_0(X0,X1)) & (r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
104. ? [X0] : (? [X1] : (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(X0,X1)) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 103]
119. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 34]
120. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 119]
123. ! [X0] : (! [X1] : ((v3_xxreal_0(X1) | v2_xxreal_0(X0) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 41]
124. ! [X0] : (! [X1] : (v3_xxreal_0(X1) | v2_xxreal_0(X0) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 123]
134. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 81]
146. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 12]
147. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 146]
148. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 15]
149. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 148]
170. ! [X0,X1] : (~v3_xxreal_0(k2_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 24]
171. ! [X0,X1] : (~v3_xxreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 170]
176. ! [X0,X1] : (v1_xreal_0(k2_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 90]
177. ! [X0,X1] : (v1_xreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 176]
192. ? [X0] : (? [X1] : (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(X0,X1)) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,X1)) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
193. ? [X1] : (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,X1)) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) => (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,sK1)) & r1_xxreal_0(k6_numbers,sK1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) [choice axiom]
194. (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,sK1)) & r1_xxreal_0(k6_numbers,sK1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 104,193,192]
228. ? [X0] : (v1_xreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0) & v1_xboole_0(X0)) => (v1_xreal_0(sK18) & v1_xxreal_0(sK18) & v1_xcmplx_0(sK18) & v1_xboole_0(sK18)) [choice axiom]
229. v1_xreal_0(sK18) & v1_xxreal_0(sK18) & v1_xcmplx_0(sK18) & v1_xboole_0(sK18) [skolemisation 59,228]
230. ? [X0] : (v1_xxreal_0(X0) & v1_xboole_0(X0)) => (v1_xxreal_0(sK19) & v1_xboole_0(sK19)) [choice axiom]
231. v1_xxreal_0(sK19) & v1_xboole_0(sK19) [skolemisation 60,230]
232. v1_xreal_0(sK0) [cnf transformation 194] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
233. v1_xreal_0(sK1) [cnf transformation 194] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
234. r1_xxreal_0(k6_numbers,sK0) [cnf transformation 194] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
235. r1_xxreal_0(k6_numbers,sK1) [cnf transformation 194] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
236. ~r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,sK1)) [cnf transformation 194] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
238. v1_xboole_0(k1_xboole_0) [cnf transformation 74] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
240. k6_numbers = k1_xboole_0 [cnf transformation 85] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
259. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 120] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
261. r1_xxreal_0(X0,X1) | v2_xxreal_0(X0) | v3_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 124] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
268. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 134] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
279. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 147] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
282. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 149] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
303. ~v3_xxreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 171] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
306. v1_xreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 177] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
341. v1_xboole_0(sK18) [cnf transformation 229] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
344. v1_xreal_0(sK18) [cnf transformation 229] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
345. v1_xboole_0(sK19) [cnf transformation 231] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
346. v1_xxreal_0(sK19) [cnf transformation 231] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
348. 1 <=> r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,sK1)) [avatar definition]
350. ~r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,sK1)) <- (~1) [avatar component clause 348] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
351. ~1 [avatar split clause 236,348]
353. 2 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
355. r1_xxreal_0(k6_numbers,sK1) <- (2) [avatar component clause 353] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
356. 2 [avatar split clause 235,353]
358. 3 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
360. r1_xxreal_0(k6_numbers,sK0) <- (3) [avatar component clause 358] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
361. 3 [avatar split clause 234,358]
363. 4 <=> v1_xreal_0(sK1) [avatar definition]
365. v1_xreal_0(sK1) <- (4) [avatar component clause 363] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
366. 4 [avatar split clause 233,363]
368. 5 <=> v1_xreal_0(sK0) [avatar definition]
370. v1_xreal_0(sK0) <- (5) [avatar component clause 368] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
371. 5 [avatar split clause 232,368]
378. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
380. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 378] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
381. 7 [avatar split clause 238,378]
388. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
390. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 388] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
391. 9 [avatar split clause 240,388]
560. 44 <=> v1_xreal_0(sK18) [avatar definition]
562. v1_xreal_0(sK18) <- (44) [avatar component clause 560] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
563. 44 [avatar split clause 344,560]
575. 47 <=> v1_xboole_0(sK18) [avatar definition]
577. v1_xboole_0(sK18) <- (47) [avatar component clause 575] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
578. 47 [avatar split clause 341,575]
580. 48 <=> v1_xxreal_0(sK19) [avatar definition]
582. v1_xxreal_0(sK19) <- (48) [avatar component clause 580] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
583. 48 [avatar split clause 346,580]
585. 49 <=> v1_xboole_0(sK19) [avatar definition]
587. v1_xboole_0(sK19) <- (49) [avatar component clause 585] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
588. 49 [avatar split clause 345,585]
589. v1_xboole_0(k6_numbers) <- (7, 9) [backward demodulation 380,390] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
591. 50 <=> v1_xboole_0(k6_numbers) [avatar definition]
593. v1_xboole_0(k6_numbers) <- (50) [avatar component clause 591] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
594. 50 | ~7 | ~9 [avatar split clause 589,388,378,591]
831. k6_numbers = X0 | ~v1_xboole_0(X0) <- (9) [forward demodulation 268,390] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
955. 103 <=> v3_xxreal_0(sK0) [avatar definition]
968. 106 <=> v3_xxreal_0(sK1) [avatar definition]
969. ~v3_xxreal_0(sK1) <- (~106) [avatar component clause 968] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1093. ~v3_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (2) [resolution 259,355] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1094. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (3) [resolution 259,360] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1100. ~v3_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 4) [subsumption resolution 1093,365] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1102. 127 <=> v1_xreal_0(k6_numbers) [avatar definition]
1104. ~v1_xreal_0(k6_numbers) <- (~127) [avatar component clause 1102] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1106. 128 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1108. v3_xxreal_0(k6_numbers) <- (128) [avatar component clause 1106] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1109. ~127 | 128 | ~106 | ~2 | ~4 [avatar split clause 1100,363,353,968,1106,1102]
1110. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 5) [subsumption resolution 1094,370] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1111. ~127 | 128 | ~103 | ~3 | ~5 [avatar split clause 1110,368,358,955,1106,1102]
1130. v2_xxreal_0(k6_numbers) | v3_xxreal_0(k2_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k2_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (~1) [resolution 261,350] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1137. 130 <=> v1_xreal_0(k2_xcmplx_0(sK0,sK1)) [avatar definition]
1139. ~v1_xreal_0(k2_xcmplx_0(sK0,sK1)) <- (~130) [avatar component clause 1137] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1141. 131 <=> v3_xxreal_0(k2_xcmplx_0(sK0,sK1)) [avatar definition]
1143. v3_xxreal_0(k2_xcmplx_0(sK0,sK1)) <- (131) [avatar component clause 1141] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1145. 132 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1147. v2_xxreal_0(k6_numbers) <- (132) [avatar component clause 1145] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1148. ~127 | ~130 | 131 | 132 | 1 [avatar split clause 1130,348,1145,1141,1137,1102]
1154. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~130) [resolution 1139,306] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1156. ~v1_xreal_0(sK0) <- (4, ~130) [subsumption resolution 1154,365] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1157. $false <- (4, 5, ~130) [subsumption resolution 1156,370] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1158. ~4 | ~5 | 130 [avatar contradiction clause 1157]
1171. 135 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1172. v1_xxreal_0(k6_numbers) <- (135) [avatar component clause 1171] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1173. ~v1_xxreal_0(k6_numbers) <- (~135) [avatar component clause 1171] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1212. ~v1_xxreal_0(X0) | ~v1_xboole_0(X0) <- (9, ~135) [superposition 1173,831] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1217. ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (131) [resolution 1143,303] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1221. v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (4, 131) [subsumption resolution 1217,365] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1222. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (4, ~106, 131) [subsumption resolution 1221,969] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1223. v3_xxreal_0(sK0) <- (4, 5, ~106, 131) [subsumption resolution 1222,370] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1226. 103 | ~4 | ~5 | 106 | ~131 [avatar split clause 1223,1141,968,368,363,955]
1259. ~v1_xboole_0(X0) | ~v1_xreal_0(X0) <- (9, ~127) [superposition 1104,831] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1399. ~v1_xboole_0(sK19) <- (9, 48, ~135) [resolution 1212,582] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
1404. $false <- (9, 48, 49, ~135) [subsumption resolution 1399,587] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
1405. ~9 | ~48 | ~49 | 135 [avatar contradiction clause 1404]
1437. ~v1_xreal_0(sK18) <- (9, 47, ~127) [resolution 1259,577] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
1440. $false <- (9, 44, 47, ~127) [subsumption resolution 1437,562] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
1441. ~9 | ~44 | ~47 | 127 [avatar contradiction clause 1440]
1456. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (128) [resolution 1108,282] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1460. ~v1_xxreal_0(k6_numbers) <- (50, 128) [subsumption resolution 1456,593] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1461. $false <- (50, 128, 135) [subsumption resolution 1460,1172] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1462. ~50 | ~128 | ~135 [avatar contradiction clause 1461]
1497. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (132) [resolution 1147,279] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1500. ~v1_xxreal_0(k6_numbers) <- (50, 132) [subsumption resolution 1497,593] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1501. $false <- (50, 132, 135) [subsumption resolution 1500,1172] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1502. ~50 | ~132 | ~135 [avatar contradiction clause 1501]
1503. $false [avatar sat refutation 351,356,361,366,371,381,391,563,578,583,588,594,1109,1111,1148,1158,1226,1405,1441,1462,1502] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1026,thDist:-1026}
% SZS output end Proof for xreal_1__t33_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 90
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 115
% Generated clauses: 692
% Active clauses: 156
% Passive clauses: 271
% Final active clauses: 138
% Final passive clauses: 101
% Main loop iterations started: 168
% 
% >>> Simplifying Inferences
% Duplicate literals: 47
% Fw subsumption resolutions: 99
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 23
% Forward subsumptions: 83
% Backward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 197
% Unit resulting resolution: 9
% Factoring: 1
% Forward superposition: 38
% Backward superposition: 23
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 41
% Split components: 95
% Sat splitting refutations: 18
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 180
% SAT solver unit clauses: 50
% SAT solver binary clauses: 58
% 
% Memory used [KB]: 5756
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% congruence closure: 0.001 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.007 s ( own 0.003 s ) 

