i: [1,275,0,2,1,1,0,0] l92_xxreal_3
a: [3,737,0,2,1,275]
s: 737
i: [1,274,0,2,1,1,0,0] l92_xxreal_3
a: [3,742,0,2,1,274]
s: 742
i: [1,273,0,2,1,1,0,0] l92_xxreal_3
a: [3,747,0,2,1,273]
s: 747
i: [1,272,0,2,1,1,0,0] l92_xxreal_3
a: [3,752,0,2,1,272]
s: 752
i: [1,277,0,2,1,0,0,255] fc2_xreal_0
a: [3,757,0,2,1,277]
s: 757
i: [1,278,0,2,1,0,0,255] fc1_xreal_0
a: [3,762,0,2,1,278]
s: 762
i: [1,279,0,2,1,0,0,255] fc1_xboole_0
a: [3,767,0,2,1,279]
s: 767
i: [1,280,0,2,1,0,0,255] fc4_xxreal_0
a: [3,772,0,2,1,280]
s: 772
i: [1,281,0,2,1,0,0,255] fc3_xxreal_0
a: [3,777,0,2,1,281]
s: 777
i: [1,282,0,2,1,0,0,255] fc1_xxreal_0
a: [3,782,0,2,1,282]
s: 782
i: [1,283,0,2,1,0,0,255] fc2_xxreal_0
a: [3,787,0,2,1,283]
s: 787
i: [1,292,0,2,1,0,0,255] fc6_ordinal1
a: [3,827,0,2,1,292]
s: 827
i: [1,291,0,2,1,0,0,255] fc6_ordinal1
a: [3,832,0,2,1,291]
s: 832
i: [1,542,0,21,9,0,0,255] d5_xxreal_3
d: [2,728,0,17,7,0,18,542]
a: [3,837,0,2,1,728]
s: 837
i: [1,498,0,2,1,0,0,2] rc2_xboole_0
a: [3,951,0,2,1,498]
s: 951
i: [1,499,0,2,1,0,0,255] rc4_ordinal1
a: [3,956,0,2,1,499]
s: 956
i: [1,500,0,2,1,0,0,255] rc1_ordinal1
a: [3,961,0,2,1,500]
s: 961
i: [1,501,0,2,1,0,0,255] rc1_xcmplx_0
a: [3,966,0,2,1,501]
s: 966
i: [1,502,0,2,1,0,0,1] rc1_xreal_0
a: [3,971,0,2,1,502]
s: 971
i: [1,503,0,2,1,0,0,2] rc1_xboole_0
a: [3,976,0,2,1,503]
s: 976
i: [1,504,0,2,1,0,0,1] rc1_xxreal_0
a: [3,981,0,2,1,504]
s: 981
i: [1,508,0,2,1,0,0,255] rc2_ordinal1
a: [3,986,0,2,1,508]
s: 986
i: [1,507,0,2,1,0,0,255] rc2_ordinal1
a: [3,991,0,2,1,507]
s: 991
i: [1,506,0,2,1,0,0,255] rc2_ordinal1
a: [3,996,0,2,1,506]
s: 996
i: [1,505,0,2,1,0,0,255] rc2_ordinal1
a: [3,1001,0,2,1,505]
s: 1001
i: [1,510,0,2,1,0,0,255] rc5_ordinal1
a: [3,1006,0,2,1,510]
s: 1006
i: [1,509,0,2,1,0,0,255] rc5_ordinal1
a: [3,1011,0,2,1,509]
s: 1011
i: [1,512,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,1016,0,2,1,512]
s: 1016
i: [1,511,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,1021,0,2,1,511]
s: 1021
i: [1,516,0,2,1,0,0,255] rc3_xreal_0
a: [3,1026,0,2,1,516]
s: 1026
i: [1,515,0,2,1,0,0,255] rc3_xreal_0
a: [3,1031,0,2,1,515]
s: 1031
i: [1,514,0,2,1,0,0,255] rc3_xreal_0
a: [3,1036,0,2,1,514]
s: 1036
i: [1,513,0,2,1,0,0,255] rc3_xreal_0
a: [3,1041,0,2,1,513]
s: 1041
i: [1,520,0,2,1,0,0,255] rc2_xreal_0
a: [3,1046,0,2,1,520]
s: 1046
i: [1,519,0,2,1,0,0,255] rc2_xreal_0
a: [3,1051,0,2,1,519]
s: 1051
i: [1,518,0,2,1,0,0,255] rc2_xreal_0
a: [3,1056,0,2,1,518]
s: 1056
i: [1,517,0,2,1,0,0,255] rc2_xreal_0
a: [3,1061,0,2,1,517]
s: 1061
i: [1,524,0,2,1,0,0,255] rc4_xreal_0
a: [3,1066,0,2,1,524]
s: 1066
i: [1,523,0,2,1,0,0,255] rc4_xreal_0
a: [3,1071,0,2,1,523]
s: 1071
i: [1,522,0,2,1,0,0,255] rc4_xreal_0
a: [3,1076,0,2,1,522]
s: 1076
i: [1,521,0,2,1,0,0,255] rc4_xreal_0
a: [3,1081,0,2,1,521]
s: 1081
i: [1,526,0,2,1,0,0,2] rc4_xxreal_0
a: [3,1086,0,2,1,526]
s: 1086
i: [1,525,0,2,1,0,0,2] rc4_xxreal_0
a: [3,1091,0,2,1,525]
s: 1091
i: [1,528,0,2,1,0,0,2] rc3_xxreal_0
a: [3,1096,0,2,1,528]
s: 1096
i: [1,527,0,2,1,0,0,2] rc3_xxreal_0
a: [3,1101,0,2,1,527]
s: 1101
i: [1,531,0,2,1,0,0,2] rc2_xxreal_3
a: [3,1106,0,2,1,531]
s: 1106
i: [1,530,0,2,1,0,0,2] rc2_xxreal_3
a: [3,1111,0,2,1,530]
s: 1111
i: [1,529,0,2,1,0,0,2] rc2_xxreal_3
a: [3,1116,0,2,1,529]
s: 1116
i: [1,533,0,2,1,0,0,2] rc2_xxreal_0
a: [3,1121,0,2,1,533]
s: 1121
i: [1,532,0,2,1,0,0,2] rc2_xxreal_0
a: [3,1126,0,2,1,532]
s: 1126
i: [1,536,0,2,1,0,0,2] rc1_xxreal_3
a: [3,1131,0,2,1,536]
s: 1131
i: [1,535,0,2,1,0,0,2] rc1_xxreal_3
a: [3,1136,0,2,1,535]
s: 1136
i: [1,534,0,2,1,0,0,2] rc1_xxreal_3
a: [3,1141,0,2,1,534]
s: 1141
i: [1,541,0,2,1,0,0,255] rc1_real_1
a: [3,1146,0,2,1,541]
s: 1146
i: [1,540,0,2,1,0,0,255] rc1_real_1
a: [3,1151,0,2,1,540]
s: 1151
i: [1,539,0,2,1,0,0,255] rc1_real_1
a: [3,1156,0,2,1,539]
s: 1156
i: [1,538,0,2,1,0,0,255] rc1_real_1
a: [3,1161,0,2,1,538]
s: 1161
i: [1,276,0,11,1,1,0,0] l92_xxreal_3
a: [3,732,0,11,1,276]
s: 732
i: [1,464,0,3,1,0,0,255] reflexivity_r1_tarski
s: 464
i: [1,284,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,792,0,3,1,284]
s: 792
i: [1,285,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,797,0,3,1,285]
s: 797
r: 767
d: [2,1168,0,2,1,2,23,767,797]
a: [3,1172,0,2,1,1168]
s: 1172
i: [1,286,0,3,1,0,0,255] d2_xxreal_0
a: [3,802,0,3,1,286]
s: 802
r: 782
r: 777
r: 757
d: [2,1174,0,2,1,2,23,782,802]
a: [3,1180,0,2,1,1174]
s: 1180
d: [2,1175,0,2,1,2,23,777,802]
a: [3,1185,0,2,1,1175]
s: 1185
d: [2,1176,0,2,1,2,23,757,802]
a: [3,1190,0,2,1,1176]
s: 1190
i: [1,287,0,3,1,0,0,255] t1_numerals
a: [3,807,0,3,1,287]
s: 807
i: [1,537,0,3,1,0,0,255] rc1_real_1
a: [3,1166,0,3,1,537]
s: 1166
d: [2,1192,0,3,1,2,22,807,797]
a: [3,1196,0,3,1,1192]
s: 1196
i: [1,294,0,4,2,0,0,255] cc6_ordinal1
s: 294
i: [1,295,0,4,2,0,0,255] cc2_xcmplx_0
s: 295
d: [2,1198,1,2,1,1,40,294,956]
a: [3,1203,1,2,1,1198]
s: 1203
i: [1,296,0,4,2,0,0,255] cc2_xreal_0
s: 296
d: [2,1199,1,2,1,1,40,294,1006]
a: [3,1208,1,2,1,1199]
s: 1208
i: [1,297,0,4,2,0,0,255] cc2_xxreal_0
s: 297
i: [1,298,0,4,2,0,0,255] cc1_ordinal1
s: 298
d: [2,1210,1,2,1,1,40,295,956]
a: [3,1215,1,2,1,1210]
s: 1215
i: [1,299,0,4,2,0,0,255] cc1_ordinal1
s: 299
d: [2,1211,1,2,1,1,40,295,1006]
a: [3,1220,1,2,1,1211]
s: 1220
i: [1,302,0,4,2,0,0,255] cc3_xreal_0
s: 302
i: [1,303,0,4,2,0,0,1] cc4_xreal_0
s: 303
d: [2,1222,1,2,1,1,40,296,956]
a: [3,1227,1,2,1,1222]
s: 1227
i: [1,304,0,4,2,0,0,255] cc9_ordinal1
s: 304
d: [2,1223,1,2,1,1,40,296,1006]
a: [3,1232,1,2,1,1223]
s: 1232
i: [1,305,0,4,2,0,0,255] cc7_ordinal1
s: 305
i: [1,306,0,4,2,0,0,255] cc3_ordinal1
s: 306
d: [2,1234,1,2,1,1,40,297,956]
a: [3,1239,1,2,1,1234]
s: 1239
i: [1,463,0,4,1,0,0,255] existence_m1_subset_1
s: 463
d: [2,1235,1,2,1,1,40,297,1006]
a: [3,1244,1,2,1,1235]
s: 1244
i: [1,288,0,4,1,0,0,255] dt_k5_numbers
a: [3,812,0,4,1,288]
s: 812
i: [1,289,0,4,1,0,0,255] dt_k6_numbers
a: [3,817,0,4,1,289]
s: 817
d: [2,1246,1,2,1,1,40,298,827]
a: [3,1254,1,2,1,1246]
s: 1254
d: [2,1394,0,4,1,2,22,812,792]
a: [3,1398,0,4,1,1394]
s: 1398
d: [2,1248,1,2,1,1,40,298,961]
a: [3,1264,1,2,1,1248]
s: 1264
d: [2,1400,0,4,1,2,22,817,792]
a: [3,1404,0,4,1,1400]
s: 1404
i: [1,307,0,5,2,0,0,255] t6_boole
s: 307
d: [2,1271,1,2,1,1,40,299,827]
a: [3,1279,1,2,1,1271]
s: 1279
i: [1,432,0,5,2,0,0,255] cc8_ordinal1
s: 432
d: [2,1273,1,2,1,1,40,299,961]
a: [3,1289,1,2,1,1273]
s: 1289
i: [1,433,0,5,2,0,0,255] cc1_xcmplx_0
s: 433
i: [1,434,0,5,2,0,0,255] cc1_xreal_0
s: 434
d: [2,1296,1,2,1,1,40,302,971]
a: [3,1304,1,2,1,1296]
s: 1304
i: [1,435,0,5,2,0,0,255] cc1_real_1
s: 435
d: [2,1306,1,2,1,1,40,303,971]
a: [3,1314,1,2,1,1306]
s: 1314
i: [1,460,0,5,2,0,0,255] t7_boole
s: 460
i: [1,290,0,5,1,0,0,255] d3_xxreal_0
a: [3,822,0,5,1,290]
s: 822
d: [2,1319,1,2,1,1,40,304,1172]
a: [3,1326,1,2,1,1319]
s: 1326
d: [2,1406,0,5,2,1,22,307,797]
s: 1406
d: [2,1320,1,2,1,1,40,304,976]
a: [3,1331,1,2,1,1320]
s: 1331
i: [1,440,0,6,3,0,0,255] cc2_ordinal1
s: 440
i: [1,442,0,6,3,0,0,2] cc7_xxreal_0
s: 442
d: [2,1321,1,2,1,1,40,304,1081]
a: [3,1336,1,2,1,1321]
s: 1336
i: [1,443,0,6,3,0,0,2] cc7_xxreal_0
s: 443
d: [2,1322,1,2,1,1,40,304,1091]
a: [3,1341,1,2,1,1322]
s: 1341
i: [1,450,0,6,3,0,0,2] cc5_xxreal_0
s: 450
i: [1,452,0,6,3,0,0,2] cc5_xxreal_0
s: 452
d: [2,1346,1,2,1,1,40,305,1172]
a: [3,1353,1,2,1,1346]
s: 1353
d: [2,1520,2,0,0,2,52,837,1353,297]
e: 1520
d: [2,1521,2,2,1,1,40,1353,297]
d: [2,1527,2,0,0,2,21,1521,837]
e: 1527
r: 837
a: [3,836,0,2,1,728]
s: 836
d: [2,1347,1,2,1,1,40,305,976]
a: [3,1358,1,2,1,1347]
s: 1358
i: [1,453,0,6,3,0,0,2] cc3_xxreal_0
s: 453
i: [1,455,0,6,3,0,0,2] cc3_xxreal_0
s: 455
d: [2,1348,1,2,1,1,40,305,1081]
a: [3,1363,1,2,1,1348]
s: 1363
i: [1,456,0,6,2,0,0,255] t1_subset
s: 456
d: [2,1349,1,2,1,1,40,305,1091]
a: [3,1368,1,2,1,1349]
s: 1368
i: [1,465,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 465
i: [1,300,0,7,3,0,0,255] cc5_ordinal1
s: 300
d: [2,1370,1,2,1,1,40,306,1172]
a: [3,1377,1,2,1,1370]
s: 1377
i: [1,301,0,7,2,0,0,255] t2_arithm
s: 301
d: [2,1371,1,2,1,1,40,306,976]
a: [3,1382,1,2,1,1371]
s: 1382
i: [1,458,0,7,2,0,0,255] t3_subset
s: 458
i: [1,459,0,7,3,0,0,1] t8_boole
s: 459
d: [2,1372,1,2,1,1,40,306,1081]
a: [3,1387,1,2,1,1372]
s: 1387
i: [1,293,0,8,3,0,0,255] cc10_ordinal1
s: 293
d: [2,1373,1,2,1,1,40,306,1091]
a: [3,1392,1,2,1,1373]
s: 1392
i: [1,437,0,8,4,0,0,2] cc4_xxreal_0
s: 437
i: [1,439,0,8,4,0,0,2] cc6_xxreal_0
s: 439
d: [2,1669,1,6,3,1,40,437,752]
a: [3,1688,1,2,1,1669]
s: 1688
r: 1688
i: [1,445,0,8,4,0,0,2] cc1_xxreal_3
s: 445
a: [3,1687,1,2,1,1669]
s: 1687
i: [1,447,0,8,4,0,0,2] cc2_xxreal_3
s: 447
i: [1,448,0,8,4,0,0,2] cc8_xxreal_0
s: 448
a: [3,1692,1,2,1,1669]
s: 1692
i: [1,457,0,8,3,0,0,255] t2_subset
s: 457
a: [3,1695,1,2,1,1669]
s: 1695
i: [1,477,0,8,3,0,0,255] fc3_xcmplx_0
s: 477
i: [1,479,0,8,3,0,0,255] fc6_xreal_0
s: 479
d: [2,1670,1,6,3,1,40,437,747]
a: [3,1701,1,2,1,1670]
s: 1701
r: 1701
i: [1,494,0,8,3,0,0,1] dt_k4_xxreal_3
s: 494
a: [3,1700,1,2,1,1670]
s: 1700
i: [1,462,0,9,3,0,0,255] t5_subset
s: 462
i: [1,461,0,10,3,0,0,255] t4_subset
s: 461
a: [3,1705,1,2,1,1670]
s: 1705
i: [1,480,0,10,4,0,0,1] fc37_xxreal_3
s: 480
a: [3,1708,1,2,1,1670]
s: 1708
i: [1,481,0,10,4,0,0,1] fc37_xxreal_3
s: 481
i: [1,308,0,11,4,0,0,255] t14_xxreal_0
s: 308
d: [2,1671,1,6,3,1,40,437,742]
a: [3,1714,1,2,1,1671]
s: 1714
r: 1714
i: [1,478,0,11,3,0,0,255] commutativity_k3_xcmplx_0
s: 478
a: [3,1713,1,2,1,1671]
s: 1713
i: [1,495,0,11,3,0,0,1] commutativity_k4_xxreal_3
s: 495
d: [2,1918,0,11,4,1,22,308,802]
s: 1918
a: [3,1718,1,2,1,1671]
s: 1718
i: [1,470,0,12,5,0,0,255] fc8_xcmplx_0
s: 470
a: [3,1721,1,2,1,1671]
s: 1721
i: [1,471,0,12,5,0,0,1] fc36_xxreal_3
s: 471
i: [1,472,0,12,5,0,0,1] fc36_xxreal_3
s: 472
d: [2,1965,1,17,3,1,44,732,495]
d: [2,1983,1,15,2,2,21,1965,742]
a: [3,1987,1,4,1,1983]
s: 1987
d: [2,2020,2,0,0,3,52,747,752,1987,494]
e: 2020
d: [2,2021,2,4,2,1,40,1987,494]
d: [2,2025,2,2,1,2,21,2021,752]
d: [2,2026,2,0,0,3,21,2025,747]
e: 2026
r: 1987
i: [1,473,0,12,5,0,0,255] fc26_xreal_0
s: 473
a: [3,1986,1,4,1,1983]
s: 1986
i: [1,474,0,12,5,0,0,255] fc23_xreal_0
s: 474
i: [1,475,0,12,5,0,0,255] fc24_xreal_0
s: 475
d: [2,2035,2,12,3,1,40,1986,445]
a: [3,2042,2,4,1,2035]
s: 2042
r: 2042
i: [1,476,0,12,5,0,0,255] fc25_xreal_0
s: 476
a: [3,2041,2,4,1,2035]
s: 2041
i: [1,482,0,12,5,0,0,1] fc29_xxreal_3
s: 482
i: [1,483,0,12,5,0,0,1] fc29_xxreal_3
s: 483
a: [3,2046,2,4,1,2035]
s: 2046
d: [2,2104,3,0,0,5,52,747,1708,752,1695,2046,483]
e: 2104
d: [2,2105,3,8,4,1,40,2046,483]
d: [2,2109,3,6,3,2,21,2105,1695]
d: [2,2110,3,4,2,3,21,2109,752]
d: [2,2111,3,2,1,4,21,2110,1708]
d: [2,2112,3,0,0,5,21,2111,747]
e: 2112
r: 2046
i: [1,484,0,12,5,0,0,1] fc27_xxreal_3
s: 484
a: [3,2050,2,4,1,2035]
s: 2050
r: 2041
r: 2050
i: [1,485,0,12,5,0,0,1] fc27_xxreal_3
s: 485
i: [1,486,0,12,5,0,0,1] fc28_xxreal_3
s: 486
s: 2042
i: [1,487,0,12,5,0,0,1] fc28_xxreal_3
s: 487
a: [3,2045,2,4,1,2035]
s: 2045
i: [1,488,0,12,5,0,0,1] fc25_xxreal_3
s: 488
i: [1,489,0,12,5,0,0,1] fc25_xxreal_3
s: 489
a: [3,2049,2,4,1,2035]
s: 2049
i: [1,490,0,12,5,0,0,1] fc24_xxreal_3
s: 490
d: [2,2036,2,12,3,1,40,1986,437]
a: [3,2055,2,4,1,2036]
s: 2055
d: [2,2185,3,0,0,5,52,1700,747,1687,752,2055,471]
e: 2185
d: [2,2186,3,0,0,3,52,2042,1986,2055,442]
e: 2186
d: [2,2187,3,8,4,1,40,2055,471]
d: [2,2197,3,6,3,2,21,2187,752]
d: [2,2198,3,4,2,3,21,2197,1687]
d: [2,2199,3,2,1,4,21,2198,747]
d: [2,2200,3,0,0,5,21,2199,1700]
e: 2200
d: [2,2189,3,8,2,1,40,2055,442]
d: [2,2202,3,4,1,2,21,2189,1986]
d: [2,2203,3,0,0,3,21,2202,2042]
e: 2203
r: 2055
i: [1,491,0,12,5,0,0,1] fc24_xxreal_3
s: 491
i: [1,492,0,12,5,0,0,1] fc26_xxreal_3
s: 492
d: [2,2073,3,8,3,1,44,2042,495]
d: [2,2077,3,6,2,2,21,2073,752]
d: [2,2078,3,4,1,3,21,2077,747]
a: [3,2082,3,4,1,2078]
s: 2082
i: [1,493,0,12,5,0,0,1] fc26_xxreal_3
s: 493
d: [2,2120,3,8,3,1,44,2050,495]
d: [2,2127,3,6,2,2,21,2120,752]
d: [2,2128,3,4,1,3,21,2127,747]
a: [3,2131,3,4,1,2128]
s: 2131
i: [1,469,0,14,4,0,0,255] existence_m2_subset_1
s: 469
i: [1,466,0,15,5,0,0,255] dt_m2_subset_1
s: 466
d: [2,2157,3,8,3,1,44,2045,495]
d: [2,2159,3,6,2,2,21,2157,752]
d: [2,2160,3,4,1,3,21,2159,747]
a: [3,2164,3,4,1,2160]
s: 2164
i: [1,467,0,15,5,0,0,255] redefinition_m2_subset_1
s: 467
d: [2,1440,1,4,2,1,40,442,976]
a: [3,1446,1,2,1,1440]
s: 1446
i: [1,468,0,15,5,0,0,255] redefinition_m2_subset_1
s: 468
i: [1,668,0,15,5,0,0,255] ie4_xxreal_3
s: 668
a: [3,1449,1,2,1,1440]
s: 1449
d: [2,2255,0,13,4,0,21,668,302]
d: [2,2256,0,11,3,0,21,2255,302]
s: 2256
d: [2,1441,1,4,2,1,40,442,1081]
d: [2,1452,1,2,1,2,21,1441,1071]
a: [3,1456,1,2,1,1452]
s: 1456
a: [3,840,0,15,6,728]
s: 840
i: [1,543,0,21,9,0,0,255] d5_xxreal_3
d: [2,727,0,17,7,0,18,543]
a: [3,844,0,15,6,727]
s: 844
d: [2,1442,1,4,2,1,40,442,1091]
d: [2,1458,1,2,1,2,21,1442,1086]
a: [3,1462,1,2,1,1458]
s: 1462
i: [1,578,0,21,9,0,0,255] d5_xxreal_3
d: [2,702,0,17,7,0,18,578]
a: [3,906,0,15,6,702]
s: 906
d: [2,1465,1,4,2,1,40,443,976]
a: [3,1471,1,2,1,1465]
s: 1471
i: [1,579,0,21,9,0,0,255] d5_xxreal_3
d: [2,701,0,17,7,0,18,579]
a: [3,910,0,15,6,701]
s: 910
i: [1,642,0,17,7,0,0,255] d5_xxreal_3
a: [3,914,0,15,6,642]
s: 914
d: [2,1466,1,4,2,1,40,443,1081]
d: [2,1473,1,2,1,2,21,1466,1071]
a: [3,1477,1,2,1,1473]
s: 1477
i: [1,643,0,17,7,0,0,255] d5_xxreal_3
a: [3,918,0,15,6,643]
s: 918
d: [2,1467,1,4,2,1,40,443,1091]
d: [2,1479,1,2,1,2,21,1467,1086]
a: [3,1483,1,2,1,1479]
s: 1483
i: [1,644,0,17,7,0,0,255] d5_xxreal_3
a: [3,922,0,15,6,644]
s: 922
i: [1,645,0,17,7,0,0,255] d5_xxreal_3
a: [3,926,0,15,6,645]
s: 926
d: [2,1485,1,4,2,1,40,452,1185]
d: [2,1490,1,2,1,2,21,1485,1180]
a: [3,1494,1,2,1,1490]
s: 1494
i: [1,650,0,17,7,0,0,255] d5_xxreal_3
a: [3,930,0,15,6,650]
s: 930
d: [2,1486,1,4,2,1,40,452,1051]
d: [2,1496,1,2,1,2,21,1486,1056]
a: [3,1500,1,2,1,1496]
s: 1500
i: [1,651,0,17,7,0,0,255] d5_xxreal_3
a: [3,934,0,15,6,651]
s: 934
i: [1,652,0,17,7,0,0,255] d5_xxreal_3
a: [3,938,0,15,6,652]
s: 938
d: [2,1487,1,4,2,1,40,452,1121]
d: [2,1502,1,2,1,2,21,1487,1126]
a: [3,1506,1,2,1,1502]
s: 1506
i: [1,653,0,17,7,0,0,255] d5_xxreal_3
a: [3,942,0,15,6,653]
s: 942
d: [2,1488,1,4,2,1,40,452,1136]
d: [2,1508,1,2,1,2,21,1488,1141]
a: [3,1512,1,2,1,1508]
s: 1512
i: [1,666,0,17,4,0,0,255] l90_xxreal_3
a: [3,946,0,15,3,666]
s: 946
d: [2,2605,0,15,6,2,22,930,802]
s: 2605
d: [2,1489,1,4,2,1,40,452,1151]
d: [2,1514,1,2,1,2,21,1489,1156]
a: [3,1518,1,2,1,1514]
s: 1518
d: [2,2606,0,15,6,2,22,934,802]
s: 2606
d: [2,1464,1,4,2,1,40,443,1172]
a: [3,1533,1,2,1,1464]
s: 1533
d: [2,2607,0,15,6,2,22,938,802]
s: 2607
d: [2,2608,0,15,6,2,22,942,802]
s: 2608
d: [2,2881,1,21,6,3,44,732,2607]
d: [2,2923,1,17,5,4,21,2881,1986]
d: [2,2924,1,15,4,5,21,2923,742]
d: [2,2925,1,11,3,6,21,2924,2042]
d: [2,2926,1,9,2,7,21,2925,1718]
a: [3,2930,1,7,1,2926]
s: 2930
i: [1,496,0,17,4,0,0,255] spc7_arithm
s: 496
d: [2,3041,2,19,6,2,44,2930,844]
d: [2,3070,2,17,5,3,21,3041,747]
a: [3,3082,2,4,1,3070]
s: 3082
d: [2,3163,3,0,0,3,52,752,742,3082,494]
e: 3163
d: [2,3164,3,4,2,1,40,3082,494]
d: [2,3176,3,2,1,2,21,3164,742]
d: [2,3177,3,0,0,3,21,3176,752]
e: 3177
d: [2,3174,3,12,6,3,44,3082,2608]
d: [2,3184,3,10,5,4,21,3174,1180]
d: [2,3185,3,8,4,5,21,3184,752]
d: [2,3186,3,6,3,6,21,3185,742]
d: [2,3187,3,4,2,7,21,3186,737]
d: [2,3188,3,2,1,8,21,3187,1692]
d: [2,3189,3,0,0,9,21,3188,1718]
e: 3189
r: 3082
i: [1,544,0,21,9,0,0,255] d5_xxreal_3
d: [2,726,0,19,8,0,18,544]
a: [3,848,0,17,7,726]
s: 848
i: [1,545,0,21,9,0,0,255] d5_xxreal_3
d: [2,725,0,19,8,0,18,545]
a: [3,852,0,17,7,725]
s: 852
a: [3,3086,2,3,1,3070]
s: 3086
i: [1,546,0,21,9,0,0,255] d5_xxreal_3
d: [2,724,0,19,8,0,18,546]
a: [3,856,0,17,7,724]
s: 856
d: [2,3045,2,19,6,2,44,2930,922]
d: [2,3088,2,17,5,3,21,3045,747]
d: [2,3089,2,15,4,4,21,3088,1705]
a: [3,3097,2,3,1,3089]
s: 3097
i: [1,547,0,21,9,0,0,255] d5_xxreal_3
d: [2,723,0,19,8,0,18,547]
a: [3,860,0,17,7,723]
s: 860
i: [1,556,0,21,9,0,0,255] d5_xxreal_3
d: [2,718,0,19,8,0,18,556]
a: [3,876,0,17,7,718]
s: 876
a: [3,3073,2,4,1,3070]
s: 3073
d: [2,3192,3,0,0,5,52,752,1692,742,1718,3073,493]
e: 3192
d: [2,3193,3,8,4,1,40,3073,493]
d: [2,3206,3,6,3,2,21,3193,1718]
d: [2,3207,3,4,2,3,21,3206,742]
d: [2,3208,3,2,1,4,21,3207,1692]
d: [2,3209,3,0,0,5,21,3208,752]
e: 3209
d: [2,3204,3,12,6,3,44,3073,2608]
d: [2,3216,3,10,5,4,21,3204,1185]
d: [2,3217,3,8,4,5,21,3216,752]
d: [2,3218,3,6,3,6,21,3217,742]
d: [2,3219,3,4,2,7,21,3218,737]
d: [2,3220,3,2,1,8,21,3219,1692]
d: [2,3221,3,0,0,9,21,3220,1718]
e: 3221
r: 1692
r: 1695
r: 1705
r: 1708
i: [1,557,0,21,9,0,0,255] d5_xxreal_3
d: [2,717,0,19,8,0,18,557]
a: [3,880,0,17,7,717]
s: 880
a: [3,1691,1,2,1,1669]
s: 1691
i: [1,564,0,21,9,0,0,255] d5_xxreal_3
d: [2,712,0,19,8,0,18,564]
a: [3,888,0,17,7,712]
s: 888
i: [1,565,0,21,9,0,0,255] d5_xxreal_3
d: [2,711,0,19,8,0,18,565]
a: [3,892,0,17,7,711]
s: 892
a: [3,1696,1,2,1,1669]
s: 1696
i: [1,646,0,19,7,0,0,255] d5_xxreal_3
s: 646
a: [3,1704,1,2,1,1670]
s: 1704
i: [1,647,0,19,7,0,0,255] d5_xxreal_3
s: 647
i: [1,648,0,19,7,0,0,255] d5_xxreal_3
s: 648
a: [3,1709,1,2,1,1670]
s: 1709
i: [1,649,0,19,7,0,0,255] d5_xxreal_3
s: 649
a: [3,3077,2,4,1,3070]
s: 3077
d: [2,3444,3,12,6,2,44,3077,918]
d: [2,3457,3,10,5,3,21,3444,772]
d: [2,3458,3,8,4,4,21,3457,752]
d: [2,3459,3,6,3,5,21,3458,742]
d: [2,3460,3,4,2,6,21,3459,737]
d: [2,3461,3,2,1,7,21,3460,1696]
d: [2,3462,3,0,0,8,21,3461,1718]
e: 3462
r: 1691
r: 1696
r: 1704
r: 1709
r: 1718
r: 1721
r: 1713
i: [1,654,0,19,7,0,0,255] d5_xxreal_3
s: 654
i: [1,655,0,19,7,0,0,255] d5_xxreal_3
s: 655
s: 1692
i: [1,656,0,19,7,0,0,255] d5_xxreal_3
s: 656
s: 1695
i: [1,657,0,19,7,0,0,255] d5_xxreal_3
s: 657
i: [1,664,0,19,6,0,0,255] d5_xxreal_3
s: 664
s: 1705
d: [2,3562,0,17,5,0,21,664,303]
d: [2,3563,0,15,4,0,21,3562,303]
s: 3563
s: 1708
i: [1,665,0,19,6,0,0,255] d5_xxreal_3
s: 665
d: [2,3585,0,17,5,0,21,665,303]
d: [2,3586,0,15,4,0,21,3585,303]
s: 3586
s: 1714
i: [1,629,0,23,9,0,0,255] d5_xxreal_3
d: [2,671,0,19,7,0,18,629]
s: 671
a: [3,1717,1,2,1,1671]
s: 1717
i: [1,628,0,23,9,0,0,255] d5_xxreal_3
d: [2,672,0,19,7,0,18,628]
s: 672
i: [1,593,0,23,9,0,0,255] d5_xxreal_3
d: [2,697,0,19,7,0,18,593]
s: 697
s: 1721
i: [1,592,0,23,9,0,0,255] d5_xxreal_3
d: [2,698,0,19,7,0,18,592]
s: 698
d: [2,1802,1,6,3,1,40,445,742]
a: [3,1827,1,2,1,1802]
s: 1827
i: [1,548,0,21,9,0,0,255] d5_xxreal_3
a: [3,864,0,19,8,548]
s: 864
i: [1,549,0,21,9,0,0,255] d5_xxreal_3
a: [3,868,0,19,8,549]
s: 868
d: [2,1921,2,2,1,1,40,1714,306]
a: [3,1931,2,2,1,1921]
s: 1931
i: [1,662,0,20,6,0,0,255] d5_xxreal_3
s: 662
d: [2,1922,2,2,1,1,40,1714,305]
a: [3,1936,2,2,1,1922]
s: 1936
d: [2,3636,0,18,5,0,21,662,303]
d: [2,3637,0,16,4,0,21,3636,303]
s: 3637
i: [1,663,0,20,6,0,0,255] d5_xxreal_3
s: 663
d: [2,1923,2,2,1,1,40,1714,304]
a: [3,1941,2,2,1,1923]
s: 1941
d: [2,3663,0,18,5,0,21,663,303]
d: [2,3664,0,16,4,0,21,3663,303]
s: 3664
d: [2,3618,2,2,1,1,40,1827,302]
a: [3,3622,2,2,1,3618]
s: 3622
i: [1,627,0,23,9,0,0,255] d5_xxreal_3
d: [2,673,0,21,8,0,18,627]
s: 673
i: [1,626,0,23,9,0,0,255] d5_xxreal_3
d: [2,674,0,21,8,0,18,626]
s: 674
d: [2,3624,3,2,1,1,40,1931,299]
a: [3,3629,3,2,1,3624]
s: 3629
i: [1,625,0,23,9,0,0,255] d5_xxreal_3
d: [2,675,0,21,8,0,18,625]
s: 675
d: [2,3625,3,2,1,1,40,1931,298]
a: [3,3634,3,2,1,3625]
s: 3634
i: [1,624,0,23,9,0,0,255] d5_xxreal_3
d: [2,676,0,21,8,0,18,624]
s: 676
i: [1,619,0,23,9,0,0,255] d5_xxreal_3
d: [2,677,0,21,8,0,18,619]
s: 677
a: [3,3081,2,4,1,3070]
s: 3081
i: [1,618,0,23,9,0,0,255] d5_xxreal_3
d: [2,678,0,21,8,0,18,618]
s: 678
a: [3,3093,2,4,1,3089]
s: 3093
i: [1,617,0,23,9,0,0,255] d5_xxreal_3
d: [2,679,0,21,8,0,18,617]
s: 679
i: [1,616,0,23,9,0,0,255] d5_xxreal_3
d: [2,680,0,21,8,0,18,616]
s: 680
d: [2,3687,3,12,3,1,40,3081,437]
d: [2,3701,3,8,2,2,21,3687,3077]
d: [2,3702,3,4,1,3,21,3701,3073]
a: [3,3706,3,4,1,3702]
s: 3706
i: [1,615,0,23,9,0,0,255] d5_xxreal_3
d: [2,681,0,21,8,0,18,615]
s: 681
d: [2,3709,3,4,1,1,40,3093,302]
a: [3,3723,3,4,1,3709]
s: 3723
i: [1,614,0,23,9,0,0,255] d5_xxreal_3
d: [2,682,0,21,8,0,18,614]
s: 682
i: [1,613,0,23,9,0,0,255] d5_xxreal_3
d: [2,683,0,21,8,0,18,613]
s: 683
d: [2,3734,4,4,1,1,40,3706,306]
a: [3,3750,4,4,1,3734]
s: 3750
i: [1,612,0,23,9,0,0,255] d5_xxreal_3
d: [2,684,0,21,8,0,18,612]
s: 684
d: [2,3735,4,4,1,1,40,3706,305]
a: [3,3755,4,4,1,3735]
s: 3755
i: [1,609,0,23,9,0,0,255] d5_xxreal_3
d: [2,685,0,21,8,0,18,609]
s: 685
i: [1,608,0,23,9,0,0,255] d5_xxreal_3
d: [2,686,0,21,8,0,18,608]
s: 686
d: [2,3736,4,4,1,1,40,3706,304]
a: [3,3760,4,4,1,3736]
s: 3760
i: [1,607,0,23,9,0,0,255] d5_xxreal_3
d: [2,687,0,21,8,0,18,607]
s: 687
d: [2,3778,5,4,1,1,40,3750,299]
a: [3,3793,5,4,1,3778]
s: 3793
i: [1,606,0,23,9,0,0,255] d5_xxreal_3
d: [2,688,0,21,8,0,18,606]
s: 688
i: [1,605,0,23,9,0,0,255] d5_xxreal_3
d: [2,689,0,21,8,0,18,605]
s: 689
d: [2,3779,5,4,1,1,40,3750,298]
a: [3,3798,5,4,1,3779]
s: 3798
i: [1,604,0,23,9,0,0,255] d5_xxreal_3
d: [2,690,0,21,8,0,18,604]
s: 690
d: [2,3039,2,15,6,3,44,2930,2607]
d: [2,3055,2,13,5,4,21,3039,752]
d: [2,3056,2,11,4,5,21,3055,742]
d: [2,3057,2,9,3,6,21,3056,1692]
d: [2,3058,2,7,2,7,21,3057,1718]
a: [3,3062,2,5,1,3058]
s: 3062
d: [2,3857,3,0,0,8,52,1185,1705,1190,1180,747,3062,2607]
e: 3857
d: [2,3868,3,13,6,3,44,3062,2607]
d: [2,3871,3,10,5,3,19,3868]
d: [2,3879,3,8,4,4,21,3871,747]
d: [2,3880,3,6,3,5,21,3879,1180]
d: [2,3881,3,4,2,6,21,3880,1190]
d: [2,3882,3,2,1,7,21,3881,1705]
d: [2,3883,3,0,0,8,21,3882,1185]
e: 3883
r: 3062
i: [1,603,0,23,9,0,0,255] d5_xxreal_3
d: [2,691,0,21,8,0,18,603]
s: 691
i: [1,602,0,23,9,0,0,255] d5_xxreal_3
d: [2,692,0,21,8,0,18,602]
s: 692
d: [2,3032,2,15,6,2,44,2930,906]
d: [2,3236,2,13,5,3,21,3032,752]
d: [2,3237,2,11,4,4,21,3236,742]
a: [3,3241,2,5,1,3237]
s: 3241
i: [1,597,0,23,9,0,0,255] d5_xxreal_3
d: [2,693,0,21,8,0,18,597]
s: 693
d: [2,3034,2,15,6,2,44,2930,918]
d: [2,3277,2,13,5,3,21,3034,752]
d: [2,3278,2,11,4,4,21,3277,742]
d: [2,3279,2,9,3,5,21,3278,737]
d: [2,3280,2,7,2,6,21,3279,1718]
a: [3,3284,2,5,1,3280]
s: 3284
i: [1,596,0,23,9,0,0,255] d5_xxreal_3
d: [2,694,0,21,8,0,18,596]
s: 694
i: [1,595,0,23,9,0,0,255] d5_xxreal_3
d: [2,695,0,21,8,0,18,595]
s: 695
d: [2,3191,3,5,2,2,44,3086,1406]
s: 3191
i: [1,594,0,23,9,0,0,255] d5_xxreal_3
d: [2,696,0,21,8,0,18,594]
s: 696
d: [2,3890,3,9,3,1,44,3241,495]
d: [2,3906,3,7,2,2,21,3890,747]
d: [2,3907,3,5,1,3,21,3906,836]
a: [3,3911,3,5,1,3907]
s: 3911
i: [1,598,0,23,9,0,0,255] d5_xxreal_3
s: 598
i: [1,599,0,23,9,0,0,255] d5_xxreal_3
s: 599
d: [2,3924,3,9,3,1,44,3284,495]
d: [2,3940,3,7,2,2,21,3924,747]
d: [2,3941,3,5,1,3,21,3940,787]
a: [3,3945,3,5,1,3941]
s: 3945
i: [1,622,0,23,9,0,0,255] d5_xxreal_3
s: 622
d: [2,2311,1,21,6,2,44,732,844]
d: [2,3494,1,17,5,3,21,2311,1986]
d: [2,3495,1,15,4,4,21,3494,742]
d: [2,3496,1,11,3,5,21,3495,2049]
a: [3,3500,1,7,1,3496]
s: 3500
i: [1,623,0,23,9,0,0,255] d5_xxreal_3
s: 623
i: [1,660,0,23,8,0,0,255] d5_xxreal_3
s: 660
d: [2,1801,1,6,3,1,40,445,747]
a: [3,1822,1,2,1,1801]
s: 1822
i: [1,661,0,29,6,0,0,255] d5_xxreal_3
s: 661
d: [2,4045,2,2,1,1,40,1822,302]
a: [3,4049,2,2,1,4045]
s: 4049
d: [2,4051,0,27,5,0,21,661,303]
d: [2,4052,0,25,4,0,21,4051,303]
s: 4052
d: [2,1439,1,4,2,1,40,442,1172]
a: [3,1538,1,2,1,1439]
s: 1538
d: [2,4007,2,15,6,2,44,3500,844]
d: [2,4029,2,13,5,3,21,4007,752]
d: [2,4030,2,11,4,4,21,4029,742]
d: [2,4031,2,9,3,5,21,4030,737]
d: [2,4032,2,7,2,6,21,4031,1721]
d: [2,4033,2,5,1,7,21,4032,1717]
a: [3,4037,2,5,1,4033]
s: 4037
d: [2,1591,1,4,2,1,40,300,1166]
a: [3,1597,1,2,1,1591]
s: 1597
d: [2,4093,3,9,3,1,44,4037,495]
d: [2,4107,3,7,2,2,21,4093,747]
d: [2,4108,3,5,1,3,21,4107,836]
a: [3,4112,3,5,1,4108]
s: 4112
d: [2,1643,1,4,2,1,40,293,1398]
a: [3,1648,1,2,1,1643]
s: 1648
d: [2,1674,1,6,3,1,40,437,981]
a: [3,1753,1,2,1,1674]
s: 1753
r: 1753
d: [2,3889,3,7,2,2,44,3241,1406]
s: 3889
a: [3,1752,1,2,1,1674]
s: 1752
d: [2,4092,3,7,2,2,44,4037,1406]
s: 4092
a: [3,1757,1,2,1,1674]
s: 1757
a: [3,1760,1,2,1,1674]
s: 1760
d: [2,3961,4,7,2,2,44,3911,1406]
s: 3961
d: [2,1842,1,5,2,1,40,457,1166]
a: [3,1864,1,2,1,1842]
s: 1864
d: [2,4117,4,7,2,2,44,4112,1406]
s: 4117
d: [2,2391,1,20,10,1,44,471,914]
d: [2,2425,1,16,8,1,18,2391]
d: [2,2432,1,14,7,1,21,2425,442]
d: [2,2433,1,12,6,1,21,2432,443]
a: [3,2443,1,2,1,2433]
s: 2443
d: [2,2394,1,20,10,1,44,485,914]
d: [2,2422,1,16,8,1,18,2394]
d: [2,2445,1,14,7,1,21,2422,452]
d: [2,2446,1,12,6,1,21,2445,452]
a: [3,2450,1,2,1,2446]
s: 2450
d: [2,1247,2,2,1,1,40,298,1203]
a: [3,1259,2,2,1,1247]
s: 1259
d: [2,2891,1,19,6,3,44,732,2607]
d: [2,2937,1,17,5,4,21,2891,747]
d: [2,2938,1,15,4,5,21,2937,752]
d: [2,2939,1,13,3,6,21,2938,737]
d: [2,2940,1,11,2,7,21,2939,1705]
d: [2,2941,1,9,1,8,21,2940,1692]
a: [3,2945,1,9,1,2941]
s: 2945
d: [2,4221,2,17,6,2,44,2945,844]
d: [2,4241,2,15,5,3,21,4221,752]
d: [2,4242,2,13,4,4,21,4241,742]
d: [2,4243,2,11,3,5,21,4242,737]
d: [2,4244,2,9,2,6,21,4243,1721]
d: [2,4245,2,7,1,7,21,4244,1717]
a: [3,4249,2,7,1,4245]
s: 4249
d: [2,1964,1,15,3,1,44,732,495]
d: [2,1976,1,13,2,2,21,1964,747]
d: [2,1977,1,11,1,3,21,1976,752]
a: [3,1981,1,11,1,1977]
s: 1981
a: [3,1991,1,11,1,1983]
s: 1991
d: [2,2219,4,8,2,1,40,2082,452]
a: [3,2223,4,4,1,2219]
s: 2223
d: [2,4342,5,0,0,3,52,752,747,2223,494]
e: 4342
d: [2,4343,5,4,2,1,40,2223,494]
d: [2,4355,5,2,1,2,21,4343,747]
d: [2,4356,5,0,0,3,21,4355,752]
e: 4356
d: [2,4353,5,12,6,3,44,2223,2608]
d: [2,4358,5,10,5,4,21,4353,1180]
d: [2,4359,5,8,4,5,21,4358,752]
d: [2,4360,5,6,3,6,21,4359,747]
d: [2,4361,5,4,2,7,21,4360,737]
d: [2,4362,5,2,1,8,21,4361,1692]
d: [2,4363,5,0,0,9,21,4362,1705]
e: 4363
r: 2223
d: [2,1408,1,3,1,0,40,432,463]
a: [3,1413,1,3,1,1408]
s: 1413
a: [3,2222,4,4,1,2219]
s: 2222
d: [2,1415,1,3,1,0,40,433,463]
a: [3,1420,1,3,1,1415]
s: 1420
d: [2,1422,1,3,1,0,40,434,463]
a: [3,1427,1,3,1,1422]
s: 1427
d: [2,4252,3,11,3,1,44,4249,495]
d: [2,4264,3,9,2,2,21,4252,747]
d: [2,4265,3,7,1,3,21,4264,836]
a: [3,4269,3,7,1,4265]
s: 4269
d: [2,1840,1,5,2,1,40,457,1196]
d: [2,1852,1,3,1,2,21,1840,832]
a: [3,1856,1,3,1,1852]
s: 1856
d: [2,1250,2,2,1,1,40,298,1208]
a: [3,1269,2,2,1,1250]
s: 1269
a: [3,1861,1,3,1,1842]
s: 1861
d: [2,1899,1,5,2,1,40,462,1398]
a: [3,1903,1,3,1,1899]
s: 1903
d: [2,4438,2,0,0,2,40,1903,1856]
e: 4438
d: [2,4436,2,0,0,2,52,1856,1903]
e: 4436
d: [2,4437,2,0,0,3,52,1196,832,1903,457]
e: 4437
r: 1903
d: [2,1272,2,2,1,1,40,299,1203]
a: [3,1284,2,2,1,1272]
s: 1284
d: [2,2233,1,11,4,1,40,466,1404]
d: [2,2236,1,7,3,2,21,2233,1398]
d: [2,2237,1,5,2,3,21,2236,832]
a: [3,2241,1,3,1,2237]
s: 2241
d: [2,1275,2,2,1,1,40,299,1208]
a: [3,1294,2,2,1,1275]
s: 1294
d: [2,1551,1,4,2,2,44,836,1406]
s: 1551
d: [2,1839,1,7,2,1,40,457,1398]
a: [3,1846,1,4,1,1839]
s: 1846
d: [2,1523,2,2,1,1,40,1353,295]
a: [3,1544,2,2,1,1523]
s: 1544
a: [3,1849,1,3,1,1839]
s: 1849
d: [2,1522,2,2,1,1,40,1353,296]
a: [3,1549,2,2,1,1522]
s: 1549
d: [2,1433,1,5,2,2,44,1196,1406]
s: 1433
d: [2,1590,1,5,2,0,40,300,463]
s: 1590
d: [2,1553,2,2,1,1,40,1358,296]
a: [3,1560,2,2,1,1553]
s: 1560
d: [2,1435,1,6,2,2,44,1404,1406]
s: 1435
d: [2,1554,2,2,1,1,40,1358,295]
a: [3,1565,2,2,1,1554]
s: 1565
d: [2,1644,1,6,2,0,40,293,463]
s: 1644
d: [2,1841,1,6,2,0,40,457,463]
s: 1841
d: [2,1574,2,2,1,1,40,1368,296]
a: [3,1580,2,2,1,1574]
s: 1580
d: [2,1906,1,6,2,1,40,461,1398]
s: 1906
d: [2,1575,2,2,1,1,40,1368,295]
a: [3,1585,2,2,1,1575]
s: 1585
d: [2,1429,1,7,3,1,44,1406,1406]
s: 1429
d: [2,1430,1,7,2,2,44,822,1406]
s: 1430
d: [2,1603,2,2,1,1,40,1377,299]
a: [3,1609,2,2,1,1603]
s: 1609
d: [2,1642,1,7,3,0,40,293,458]
s: 1642
d: [2,1604,2,2,1,1,40,1377,298]
a: [3,1614,2,2,1,1604]
s: 1614
d: [2,1900,1,7,2,0,40,462,463]
s: 1900
d: [2,1943,1,11,4,0,44,478,301]
d: [2,1956,1,9,3,0,18,1943]
a: [3,1959,1,7,2,1956]
s: 1959
d: [2,1617,2,2,1,1,40,1382,299]
a: [3,1622,2,2,1,1617]
s: 1622
d: [2,2248,1,11,4,1,40,468,1398]
d: [2,2250,1,9,3,2,21,2248,832]
a: [3,2253,1,7,2,2250]
s: 2253
d: [2,1618,2,2,1,1,40,1382,298]
a: [3,1627,2,2,1,1618]
s: 1627
d: [2,2261,1,11,4,0,44,2256,301]
d: [2,2279,1,9,3,0,21,2261,302]
a: [3,2282,1,7,2,2279]
s: 2282
d: [2,4543,3,0,0,3,52,1822,4037,2282]
e: 4543
d: [2,4583,3,5,2,2,44,4037,2282]
d: [2,4588,3,2,1,2,19,4583]
d: [2,4618,3,0,0,3,21,4588,1822]
e: 4618
d: [2,4585,4,7,3,3,44,4092,2282]
d: [2,4587,4,4,2,3,19,4585]
d: [2,4620,4,2,1,4,21,4587,1172]
d: [2,4621,4,0,0,5,21,4620,1822]
e: 4621
r: 2049
r: 2131
r: 3500
r: 4092
r: 4037
a: [3,3499,1,7,1,3496]
s: 3499
r: 1991
r: 1981
r: 2945
r: 732
s: 2050
d: [2,4765,3,12,6,3,44,2050,2607]
d: [2,4774,3,10,5,4,21,4765,1190]
d: [2,4775,3,8,4,5,21,4774,747]
d: [2,4776,3,6,3,6,21,4775,752]
d: [2,4777,3,4,2,7,21,4776,737]
d: [2,4778,3,2,1,8,21,4777,1705]
d: [2,4779,3,0,0,9,21,4778,1692]
e: 4779
r: 1705
r: 1708
r: 2042
r: 2045
r: 2082
r: 2164
r: 1700
d: [2,4648,0,7,1,2,23,732,3499]
a: [3,4705,0,7,1,4648]
s: 4705
s: 1701
s: 1704
s: 1708
d: [2,1874,2,2,1,1,40,1701,306]
a: [3,1884,2,2,1,1874]
s: 1884
d: [2,4645,1,7,1,2,23,1991,3499]
a: [3,4694,1,7,1,4645]
s: 4694
d: [2,1875,2,2,1,1,40,1701,305]
a: [3,1889,2,2,1,1875]
s: 1889
d: [2,4646,1,7,1,2,23,1981,3499]
a: [3,4699,1,7,1,4646]
s: 4699
d: [2,1629,1,8,3,0,40,458,300]
s: 1629
d: [2,1876,2,2,1,1,40,1701,304]
a: [3,1894,2,2,1,1876]
s: 1894
d: [2,1869,1,8,3,0,40,479,303]
s: 1869
d: [2,4926,3,2,1,1,40,1884,299]
a: [3,4931,3,2,1,4926]
s: 4931
d: [2,1870,1,8,3,0,40,479,302]
s: 1870
d: [2,1898,1,8,3,0,40,462,458]
s: 1898
d: [2,4927,3,2,1,1,40,1884,298]
a: [3,4936,3,2,1,4927]
s: 4936
d: [2,1907,1,8,2,0,40,461,463]
s: 1907
s: 2041
d: [2,2000,1,10,4,1,40,1918,460]
a: [3,2003,1,8,3,2000]
s: 2003
r: 1918
d: [2,2273,1,12,5,0,44,479,2256]
d: [2,2274,1,8,3,0,18,2273]
s: 2274
d: [2,5071,2,7,3,2,44,2443,2003]
d: [2,5078,2,5,2,4,21,5071,1551]
d: [2,5079,2,2,1,5,21,5078,3191]
s: 5079
r: 1021
r: 1011
r: 1001
r: 1752
r: 951
r: 1687
r: 2443
r: 470
r: 832
r: 1864
r: 1849
r: 471
r: 4117
r: 3961
r: 3889
r: 3191
r: 1900
r: 1898
r: 1551
r: 1435
r: 1433
r: 1430
r: 1406
r: 480
r: 462
r: 460
r: 459
r: 443
r: 442
r: 306
r: 305
r: 304
d: [2,5113,3,0,0,6,40,5079,1091]
e: 5113
d: [2,5112,3,0,0,6,40,5079,1081]
e: 5112
d: [2,5111,3,0,0,6,40,5079,976]
e: 5111
d: [2,5110,3,0,0,6,40,5079,1714]
e: 5110
d: [2,5109,3,0,0,6,40,5079,1701]
e: 5109
d: [2,5108,3,0,0,6,40,5079,1172]
e: 5108
d: [2,5107,4,0,0,6,40,5079,3706]
e: 5107
d: [2,5099,4,0,0,6,52,3706,5079]
e: 5099
d: [2,5100,3,0,0,8,52,3081,3073,3077,5079,437]
e: 5100
d: [2,5101,3,0,0,7,52,836,1538,1533,5079,437]
e: 5101
d: [2,5102,3,0,0,8,52,747,1704,1708,5079,437]
e: 5102
d: [2,5103,3,0,0,8,52,742,1717,1721,5079,437]
e: 5103
d: [2,5104,3,0,0,8,52,1449,1446,1471,5079,437]
e: 5104
d: [2,5105,3,0,0,8,52,1071,1456,1477,5079,437]
e: 5105
d: [2,5106,3,0,0,8,52,1086,1462,1483,5079,437]
e: 5106
r: 5079
r: 3086
s: 832
s: 1687
s: 951
s: 1001
s: 3086
s: 1011
s: 2055
s: 1021
s: 304
a: [3,2081,3,4,1,2078]
s: 2081
s: 305
d: [2,2119,3,4,1,1,40,2050,302]
a: [3,2125,3,4,1,2119]
s: 2125
s: 306
s: 460
a: [3,2132,3,4,1,2128]
s: 2132
s: 1406
s: 2164
s: 442
s: 443
s: 2045
s: 459
d: [2,5179,3,4,1,1,40,304,2055]
a: [3,5189,3,4,1,5179]
s: 5189
s: 462
s: 480
d: [2,5204,3,4,1,1,40,305,2055]
a: [3,5214,3,4,1,5204]
s: 5214
s: 470
d: [2,5239,3,4,1,1,40,306,2055]
a: [3,5249,3,4,1,5239]
s: 5249
s: 471
s: 1752
d: [2,5216,4,8,3,1,44,2125,495]
d: [2,5228,4,6,2,2,21,5216,747]
d: [2,5229,4,4,1,3,21,5228,752]
a: [3,5233,4,4,1,5229]
s: 5233
s: 1864
d: [2,5340,4,8,3,1,44,5189,495]
d: [2,5352,4,6,2,2,21,5340,747]
d: [2,5353,4,4,1,3,21,5352,752]
a: [3,5357,4,4,1,5353]
s: 5357
s: 1849
s: 1551
d: [2,5405,4,8,3,1,44,5214,495]
d: [2,5418,4,6,2,2,21,5405,747]
d: [2,5419,4,4,1,3,21,5418,752]
a: [3,5423,4,4,1,5419]
s: 5423
d: [2,2458,1,20,10,1,44,471,918]
d: [2,2492,1,16,8,1,18,2458]
d: [2,2499,1,14,7,1,21,2492,442]
d: [2,2500,1,12,6,1,21,2499,443]
a: [3,2506,1,4,2,2500]
s: 2506
r: 2608
r: 2607
r: 926
r: 922
r: 918
r: 914
r: 493
r: 491
r: 452
r: 442
d: [2,5579,2,0,0,3,52,1180,1185,2506]
e: 5579
d: [2,5580,2,0,0,3,52,752,1692,2506]
e: 5580
d: [2,5581,2,0,0,3,52,981,1757,2506]
e: 5581
d: [2,5582,2,0,0,3,52,1056,1051,2506]
e: 5582
d: [2,5583,2,0,0,3,52,1126,1121,2506]
e: 5583
d: [2,5584,2,0,0,3,52,1141,1136,2506]
e: 5584
d: [2,5585,2,0,0,3,52,1156,1151,2506]
e: 5585
d: [2,5587,2,2,1,2,40,2506,1185]
d: [2,5602,2,0,0,3,21,5587,1180]
e: 5602
d: [2,5588,2,2,1,2,40,2506,1692]
d: [2,5604,2,0,0,3,21,5588,752]
e: 5604
d: [2,5589,2,2,1,2,40,2506,1757]
d: [2,5606,2,0,0,3,21,5589,981]
e: 5606
d: [2,5590,2,2,1,2,40,2506,1051]
d: [2,5608,2,0,0,3,21,5590,1056]
e: 5608
d: [2,5591,2,2,1,2,40,2506,1121]
d: [2,5610,2,0,0,3,21,5591,1126]
e: 5610
d: [2,5592,2,2,1,2,40,2506,1136]
d: [2,5612,2,0,0,3,21,5592,1141]
e: 5612
d: [2,5593,2,2,1,2,40,2506,1151]
d: [2,5614,2,0,0,3,21,5593,1156]
e: 5614
r: 2506
d: [2,5445,4,4,1,1,40,5249,299]
a: [3,5462,4,4,1,5445]
s: 5462
d: [2,5616,0,4,2,2,21,442,1551]
s: 5616
s: 452
d: [2,5446,4,4,1,1,40,5249,298]
a: [3,5467,4,4,1,5446]
s: 5467
s: 491
d: [2,5447,4,8,3,1,44,5249,495]
d: [2,5469,4,6,2,2,21,5447,747]
d: [2,5470,4,4,1,3,21,5469,752]
a: [3,5474,4,4,1,5470]
s: 5474
s: 493
s: 914
d: [2,4209,5,15,7,3,44,4117,844]
d: [2,4808,5,13,6,4,21,4209,747]
d: [2,4809,5,11,5,5,21,4808,1551]
d: [2,4810,5,8,4,5,21,4809,1406]
a: [3,4813,5,4,2,4810]
s: 4813
s: 918
d: [2,4390,5,12,3,1,40,2222,437]
a: [3,4854,5,4,1,4390]
s: 4854
s: 922
s: 926
d: [2,5618,5,8,3,1,44,5462,495]
d: [2,5624,5,6,2,2,21,5618,747]
d: [2,5625,5,4,1,3,21,5624,752]
a: [3,5629,5,4,1,5625]
s: 5629
s: 2607
d: [2,5653,5,8,3,1,44,5467,495]
d: [2,5659,5,6,2,2,21,5653,747]
d: [2,5660,5,4,1,3,21,5659,752]
a: [3,5664,5,4,1,5660]
s: 5664
s: 2608
d: [2,5273,1,4,2,2,44,836,1406]
s: 5273
a: [3,4036,2,5,1,4033]
s: 4036
r: 4249
d: [2,6444,3,9,3,1,44,4036,495]
d: [2,6481,3,4,2,2,21,6444,4112]
d: [2,6482,3,2,1,3,21,6481,747]
d: [2,6483,3,0,0,4,21,6482,836]
e: 6483
d: [2,6445,3,9,3,1,44,4036,495]
d: [2,6485,3,4,2,2,21,6445,4112]
d: [2,6486,3,2,1,3,21,6485,836]
d: [2,6487,3,0,0,4,21,6486,747]
e: 6487
d: [2,6464,3,9,3,1,44,495,4036]
d: [2,6489,3,4,2,2,21,6464,4112]
d: [2,6490,3,2,1,3,21,6489,747]
d: [2,6491,3,0,0,4,21,6490,836]
e: 6491
d: [2,6465,3,9,3,1,44,495,4036]
d: [2,6493,3,4,2,2,21,6465,4112]
d: [2,6494,3,2,1,3,21,6493,836]
d: [2,6495,3,0,0,4,21,6494,747]
e: 6495
r: 4112
s: 1433
d: [2,4307,2,19,6,2,44,1991,906]
d: [2,4790,2,17,5,3,21,4307,747]
d: [2,4791,2,15,4,4,21,4790,752]
d: [2,4792,2,13,3,5,21,4791,737]
d: [2,4793,2,9,3,6,22,4792,3499]
a: [3,4797,2,5,1,4793]
s: 4797
d: [2,6553,3,0,0,3,52,1827,4797,2282]
e: 6553
d: [2,6555,3,5,2,2,44,4797,2282]
d: [2,6570,3,2,1,2,19,6555]
d: [2,6572,3,0,0,3,21,6570,1827]
e: 6572
r: 1692
r: 1695
r: 1701
r: 1704
r: 2045
r: 2055
r: 2164
r: 4797
r: 1708
r: 3284
r: 3945
r: 4854
r: 1714
r: 1717
r: 1827
r: 2930
r: 3077
r: 1936
r: 3093
r: 3706
r: 3755
r: 1597
r: 1846
r: 1849
r: 1894
r: 1931
r: 1941
r: 2003
r: 2050
r: 2132
r: 3073
r: 3241
r: 3499
r: 3629
r: 3634
r: 3723
r: 3760
r: 4705
r: 5189
r: 5214
r: 5249
r: 5357
r: 5423
r: 5462
r: 5474
r: 5629
s: 732
s: 1918
s: 1692
s: 1695
s: 1700
s: 1704
s: 1709
s: 1713
s: 1718
s: 1827
a: [3,2929,1,7,1,2926]
s: 2929
r: 732
r: 2929
s: 732
a: [3,1930,2,2,1,1921]
s: 1930
s: 2930
d: [2,5891,5,12,6,2,44,5474,918]
d: [2,6602,5,10,5,3,21,5891,752]
d: [2,6603,5,8,4,4,21,6602,747]
d: [2,6604,5,6,3,5,21,6603,737]
a: [3,6608,5,2,1,6604]
s: 6608
r: 6608
s: 3500
d: [2,2576,1,23,6,2,44,732,926]
d: [2,3538,1,19,5,3,21,2576,1986]
d: [2,3539,1,17,4,4,21,3538,742]
d: [2,3540,1,13,3,5,21,3539,2049]
d: [2,3541,1,9,2,6,21,3540,2042]
a: [3,3545,1,7,1,3541]
s: 3545
r: 1713
r: 1718
r: 1822
r: 1930
r: 4036
r: 4694
r: 4699
r: 1889
r: 3545
s: 1714
a: [3,4704,0,7,1,4648]
s: 4704
r: 732
s: 1717
a: [3,3544,1,7,1,3541]
s: 3544
r: 3500
a: [3,1821,1,2,1,1801]
s: 1821
a: [3,1888,2,2,1,1875]
s: 1888
a: [3,4693,1,7,1,4645]
s: 4693
s: 1931
a: [3,4698,1,7,1,4646]
s: 4698
a: [3,1596,1,2,1,1591]
s: 1596
s: 1936
a: [3,1601,1,2,1,1591]
s: 1601
s: 1941
s: 2443
s: 2945
s: 3629
d: [2,2319,1,19,6,2,44,732,844]
d: [2,3251,1,17,5,3,21,2319,747]
d: [2,3252,1,15,4,4,21,3251,752]
a: [3,3256,1,9,1,3252]
s: 3256
s: 3634
d: [2,8619,1,5,1,2,22,3256,3544]
a: [3,8623,1,5,1,8619]
s: 8623
d: [2,2584,1,19,6,2,44,732,926]
d: [2,3290,1,17,5,3,21,2584,747]
d: [2,3291,1,15,4,4,21,3290,752]
d: [2,3292,1,13,3,5,21,3291,1705]
a: [3,3296,1,9,1,3292]
s: 3296
a: [3,6607,5,2,1,6604]
s: 6607
d: [2,8640,1,5,1,2,22,3296,3544]
a: [3,8644,1,5,1,8640]
s: 8644
d: [2,5890,5,12,6,2,44,5423,918]
d: [2,6610,5,10,5,3,21,5890,752]
d: [2,6611,5,8,4,4,21,6610,747]
d: [2,6612,5,6,3,5,21,6611,737]
a: [3,6616,5,2,1,6612]
s: 6616
d: [2,8666,6,0,0,2,52,762,6616,296]
e: 8666
d: [2,8667,6,0,0,2,52,6607,6616,294]
e: 8667
d: [2,8669,6,2,1,1,40,6616,296]
d: [2,8674,6,0,0,2,21,8669,762]
e: 8674
d: [2,8671,6,2,1,1,40,6616,294]
d: [2,8676,6,0,0,2,21,8671,6607]
e: 8676
r: 6616
s: 1981
s: 1991
a: [3,6615,5,2,1,6612]
s: 6615
a: [3,1850,1,3,1,1839]
s: 1850
r: 1850
d: [2,5889,5,12,6,2,44,5357,918]
d: [2,6618,5,10,5,3,21,5889,752]
d: [2,6619,5,8,4,4,21,6618,747]
d: [2,6620,5,6,3,5,21,6619,737]
a: [3,6624,5,2,1,6620]
s: 6624
s: 1849
s: 1846
d: [2,5888,5,12,6,2,44,5233,918]
d: [2,6626,5,10,5,3,21,5888,752]
d: [2,6627,5,8,4,4,21,6626,747]
d: [2,6628,5,6,3,5,21,6627,737]
a: [3,6632,5,2,1,6628]
s: 6632
d: [2,5276,1,5,2,2,44,1196,1406]
s: 5276
d: [2,6371,5,12,6,3,44,5233,2608]
d: [2,6721,5,10,5,4,21,6371,752]
d: [2,6722,5,8,4,5,21,6721,747]
d: [2,6723,5,6,3,6,21,6722,737]
a: [3,6727,5,2,1,6723]
s: 6727
s: 1435
a: [3,2439,1,6,3,2433]
s: 2439
r: 2608
r: 2607
r: 926
r: 914
d: [2,8735,2,0,0,4,52,1190,1180,1185,2439]
e: 8735
d: [2,8736,2,0,0,4,52,737,752,1692,2439]
e: 8736
d: [2,8737,2,0,0,4,52,1131,1141,1136,2439]
e: 8737
d: [2,8740,2,4,2,2,40,2439,1185]
d: [2,8752,2,2,1,3,21,8740,1180]
d: [2,8753,2,0,0,4,21,8752,1190]
e: 8753
d: [2,8741,2,4,2,2,40,2439,1692]
d: [2,8755,2,2,1,3,21,8741,752]
d: [2,8756,2,0,0,4,21,8755,737]
e: 8756
d: [2,8745,2,4,2,2,40,2439,1136]
d: [2,8760,2,2,1,3,21,8745,1141]
d: [2,8761,2,0,0,4,21,8760,1131]
e: 8761
r: 2439
d: [2,6316,6,12,6,2,44,5664,918]
d: [2,6582,6,10,5,3,21,6316,752]
d: [2,6583,6,8,4,4,21,6582,747]
d: [2,6584,6,6,3,5,21,6583,737]
a: [3,6588,6,2,1,6584]
s: 6588
s: 914
d: [2,6376,6,12,6,3,44,5664,2608]
d: [2,6687,6,10,5,4,21,6376,752]
d: [2,6688,6,8,4,5,21,6687,747]
d: [2,6689,6,6,3,6,21,6688,737]
a: [3,6693,6,2,1,6689]
s: 6693
s: 926
s: 2607
d: [2,6375,6,12,6,3,44,5629,2608]
d: [2,6695,6,10,5,4,21,6375,752]
d: [2,6696,6,8,4,5,21,6695,747]
d: [2,6697,6,6,3,6,21,6696,737]
a: [3,6701,6,2,1,6697]
s: 6701
s: 2608
d: [2,6157,2,17,6,2,44,4694,926]
d: [2,7366,2,15,5,3,21,6157,742]
d: [2,7367,2,11,4,4,21,7366,1986]
a: [3,7371,2,3,1,7367]
s: 7371
a: [3,2503,1,6,3,2500]
s: 2503
r: 2606
r: 2605
r: 922
r: 918
d: [2,9151,2,0,0,4,52,762,787,772,2503]
e: 9151
d: [2,9152,2,0,0,4,52,1821,747,1709,2503]
e: 9152
d: [2,9153,2,0,0,4,52,1106,1116,1111,2503]
e: 9153
d: [2,9155,2,4,2,2,40,2503,772]
d: [2,9164,2,2,1,3,21,9155,787]
d: [2,9165,2,0,0,4,21,9164,762]
e: 9165
d: [2,9156,2,4,2,2,40,2503,1709]
d: [2,9167,2,2,1,3,21,9156,747]
d: [2,9168,2,0,0,4,21,9167,1821]
e: 9168
d: [2,9159,2,4,2,2,40,2503,1111]
d: [2,9171,2,2,1,3,21,9159,1116]
d: [2,9172,2,0,0,4,21,9171,1106]
e: 9172
r: 2503
s: 918
s: 2046
s: 922
s: 2049
s: 2605
s: 2606
a: [3,2054,2,4,1,2036]
s: 2054
d: [2,5280,1,6,2,2,44,1404,1406]
s: 5280
a: [3,3074,2,4,1,3070]
s: 3074
d: [2,9651,3,12,6,2,44,3074,844]
d: [2,9663,3,10,5,3,21,9651,1538]
d: [2,9664,3,8,4,4,21,9663,752]
d: [2,9665,3,6,3,5,21,9664,742]
d: [2,9666,3,4,2,6,21,9665,737]
d: [2,9667,3,2,1,7,21,9666,1721]
d: [2,9668,3,0,0,8,21,9667,1717]
e: 9668
r: 1714
r: 1721
r: 2930
r: 3074
r: 3544
d: [2,7774,0,7,1,2,23,732,2929]
a: [3,7815,0,7,1,7774]
s: 7815
s: 1714
s: 1721
d: [2,9849,1,15,6,2,44,7815,914]
d: [2,9884,1,13,5,3,21,9849,747]
d: [2,9885,1,11,4,4,21,9884,752]
d: [2,9886,1,9,3,5,21,9885,737]
d: [2,9887,1,7,2,6,21,9886,1709]
d: [2,9888,1,5,1,7,21,9887,1692]
a: [3,9892,1,5,1,9888]
s: 9892
s: 3073
s: 2929
s: 3500
a: [3,3078,2,4,1,3070]
s: 3078
d: [2,10011,3,0,0,5,52,752,1695,742,1721,3078,483]
e: 10011
d: [2,10013,3,8,4,1,40,3078,483]
d: [2,10025,3,6,3,2,21,10013,1721]
d: [2,10026,3,4,2,3,21,10025,742]
d: [2,10027,3,2,1,4,21,10026,1695]
d: [2,10028,3,0,0,5,21,10027,752]
e: 10028
d: [2,10014,3,12,6,2,44,3078,844]
d: [2,10030,3,10,5,3,21,10014,1533]
d: [2,10031,3,8,4,4,21,10030,752]
d: [2,10032,3,6,3,5,21,10031,742]
d: [2,10033,3,4,2,6,21,10032,737]
d: [2,10034,3,2,1,7,21,10033,1721]
d: [2,10035,3,0,0,8,21,10034,1717]
e: 10035
r: 1704
r: 1709
r: 1714
r: 1721
r: 2041
r: 2046
r: 2081
r: 3750
r: 4693
r: 4698
r: 4704
r: 2929
r: 6607
r: 9892
s: 4705
s: 1705
s: 732
s: 1708
s: 1713
a: [3,1722,1,2,1,1671]
s: 1722
s: 6608
a: [3,9891,1,5,1,9888]
s: 9891
r: 8644
s: 2930
d: [2,6189,6,12,6,2,44,5629,918]
d: [2,6590,6,10,5,3,21,6189,752]
d: [2,6591,6,8,4,4,21,6590,747]
d: [2,6592,6,6,3,5,21,6591,737]
a: [3,6596,6,2,1,6592]
s: 6596
s: 3500
s: 2042
s: 3545
d: [2,10492,3,0,0,7,52,747,1705,1821,3078,3081,3545,926]
e: 10492
d: [2,10508,2,17,6,2,44,3545,926]
d: [2,10513,2,14,5,2,19,10508]
d: [2,10527,2,12,4,3,21,10513,747]
d: [2,10528,2,8,3,4,21,10527,3081]
d: [2,10529,2,6,2,5,21,10528,1821]
d: [2,10530,2,4,1,6,21,10529,1705]
d: [2,10531,2,0,0,7,21,10530,3078]
e: 10531
r: 1821
s: 1822
s: 2045
s: 4694
s: 3093
d: [2,10627,3,12,6,2,44,3093,926]
d: [2,10633,3,10,5,3,21,10627,762]
d: [2,10634,3,8,4,4,21,10633,752]
d: [2,10635,3,6,3,5,21,10634,742]
d: [2,10636,3,4,2,6,21,10635,737]
d: [2,10637,3,2,1,7,21,10636,1692]
d: [2,10638,3,0,0,8,21,10637,1722]
e: 10638
r: 1692
r: 1695
r: 1705
r: 1708
r: 3073
r: 3078
d: [2,1407,1,2,1,1,40,432,1196]
d: [2,1773,2,4,2,1,40,1688,443]
d: [2,1778,2,2,1,2,21,1773,752]
d: [2,1774,2,4,2,1,40,1688,442]
d: [2,1780,2,2,1,2,21,1774,752]
d: [2,1800,1,6,3,1,40,445,752]
d: [2,1817,1,4,2,2,21,1800,737]
d: [2,2264,1,11,4,0,44,301,2256]
d: [2,2288,1,9,3,0,21,2264,302]
d: [2,2973,1,23,6,3,44,732,2608]
d: [2,3015,1,19,5,4,21,2973,1986]
d: [2,3016,1,17,4,5,21,3015,742]
d: [2,3017,1,13,3,6,21,3016,2049]
d: [2,3018,1,9,2,7,21,3017,2042]
d: [2,3019,1,7,1,8,21,3018,1718]
d: [2,3049,2,19,6,3,44,2930,2607]
d: [2,3052,2,16,5,3,19,3049]
d: [2,3102,2,14,4,4,21,3052,747]
d: [2,3103,2,12,3,5,21,3102,1705]
d: [2,3050,2,17,6,3,44,2930,2608]
d: [2,3051,2,14,5,3,19,3050]
d: [2,3105,2,12,4,4,21,3051,747]
d: [2,3106,2,10,3,5,21,3105,1705]
d: [2,2315,3,12,6,2,44,2049,844]
d: [2,3261,3,10,5,3,21,2315,747]
d: [2,3262,3,8,4,4,21,3261,752]
d: [2,2476,3,12,6,2,44,2045,918]
d: [2,3420,3,10,5,3,21,2476,772]
d: [2,3421,3,8,4,4,21,3420,747]
d: [2,3422,3,6,3,5,21,3421,752]
d: [2,2243,4,8,4,1,40,2164,491]
d: [2,3437,4,6,3,2,21,2243,747]
d: [2,3438,4,4,2,3,21,3437,752]
d: [2,3440,3,8,4,1,40,3077,491]
d: [2,3502,3,6,3,2,21,3440,742]
d: [2,3503,3,4,2,3,21,3502,752]
d: [2,3194,3,8,4,1,40,3073,489]
d: [2,3528,3,6,3,2,21,3194,742]
d: [2,3529,3,4,2,3,21,3528,752]
d: [2,4917,2,2,1,1,40,1701,305]
d: [2,5318,3,8,2,1,40,443,2055]
d: [2,5325,3,4,1,2,21,5318,1986]
d: [2,6269,3,12,6,3,44,2050,2607]
d: [2,6782,3,10,5,4,21,6269,1190]
d: [2,6783,3,8,4,5,21,6782,747]
d: [2,6784,3,6,3,6,21,6783,752]
d: [2,6785,3,4,2,7,21,6784,737]
d: [2,6039,3,12,6,2,44,2050,922]
d: [2,6824,3,10,5,3,21,6039,762]
d: [2,6825,3,8,4,4,21,6824,747]
d: [2,6826,3,6,3,5,21,6825,752]
d: [2,6827,3,4,2,6,21,6826,737]
d: [2,5047,3,8,4,1,40,2041,493]
d: [2,6900,3,6,3,2,21,5047,752]
d: [2,6901,3,4,2,3,21,6900,747]
d: [2,4763,3,12,6,3,44,2050,2605]
d: [2,6921,3,10,5,4,21,4763,1190]
d: [2,6922,3,8,4,5,21,6921,747]
d: [2,6923,3,6,3,6,21,6922,752]
d: [2,6924,3,4,2,7,21,6923,737]
d: [2,5048,3,8,4,1,40,2041,489]
d: [2,6939,3,6,3,2,21,5048,752]
d: [2,6940,3,4,2,3,21,6939,747]
d: [2,5792,3,12,6,2,44,2050,914]
d: [2,7049,3,10,5,3,21,5792,762]
d: [2,7050,3,8,4,4,21,7049,747]
d: [2,7051,3,6,3,5,21,7050,752]
d: [2,7052,3,4,2,6,21,7051,737]
d: [2,5192,4,8,4,1,40,2081,489]
d: [2,7113,4,6,3,2,21,5192,747]
d: [2,7114,4,4,2,3,21,7113,752]
d: [2,5896,3,12,6,2,44,3093,918]
d: [2,7192,3,10,5,3,21,5896,762]
d: [2,7193,3,8,4,4,21,7192,752]
d: [2,7194,3,6,3,5,21,7193,742]
d: [2,7195,3,4,2,6,21,7194,737]
d: [2,5701,3,8,4,1,40,493,3073]
d: [2,7278,3,6,3,2,21,5701,742]
d: [2,7279,3,4,2,3,21,7278,752]
d: [2,3719,3,12,6,3,44,3093,2608]
d: [2,7326,3,10,5,4,21,3719,1190]
d: [2,7327,3,8,4,5,21,7326,752]
d: [2,7328,3,6,3,6,21,7327,742]
d: [2,7329,3,4,2,7,21,7328,737]
d: [2,4320,2,23,6,2,44,1991,922]
d: [2,7622,2,21,5,3,21,4320,742]
d: [2,7623,2,17,4,4,21,7622,1986]
d: [2,4289,2,23,6,2,44,1981,918]
d: [2,7642,2,19,5,3,21,4289,2222]
d: [2,7643,2,17,4,4,21,7642,742]
d: [2,3686,3,12,3,1,40,3081,445]
d: [2,4316,2,23,6,2,44,1991,844]
d: [2,7666,2,21,5,3,21,4316,742]
d: [2,7667,2,17,4,4,21,7666,1986]
d: [2,7668,2,13,3,5,21,7667,2041]
d: [2,8004,2,19,6,2,44,3545,914]
d: [2,8015,2,16,5,2,19,8004]
d: [2,8026,2,14,4,3,21,8015,747]
d: [2,8027,2,10,3,4,21,8026,3081]
d: [2,8028,2,8,2,5,21,8027,1709]
d: [2,7970,2,19,6,2,44,3500,844]
d: [2,7981,2,16,5,2,19,7970]
d: [2,8063,2,14,4,3,21,7981,747]
d: [2,8064,2,10,3,4,21,8063,3081]
d: [2,4979,3,2,1,1,40,1889,296]
d: [2,7893,1,23,6,2,44,732,918]
d: [2,7918,1,19,5,3,21,7893,1986]
d: [2,7919,1,17,4,4,21,7918,742]
d: [2,9388,3,8,3,1,44,2049,495]
d: [2,9400,3,6,2,2,21,9388,752]
d: [2,9401,3,4,1,3,21,9400,747]
d: [2,9461,2,19,6,3,44,2930,2605]
d: [2,9471,2,16,5,3,19,9461]
d: [2,9511,2,14,4,4,21,9471,747]
d: [2,9512,2,10,3,5,21,9511,3081]
d: [2,9513,2,8,2,6,21,9512,1709]
d: [2,9573,2,17,6,3,44,2930,2606]
d: [2,9583,2,14,5,3,19,9573]
d: [2,9625,2,12,4,4,21,9583,747]
d: [2,9626,2,8,3,5,21,9625,3081]
d: [2,9627,2,6,2,6,21,9626,1821]
d: [2,9628,2,4,1,7,21,9627,1709]
d: [2,7960,2,15,6,2,44,3500,844]
d: [2,8041,2,13,5,3,21,7960,752]
d: [2,8042,2,11,4,4,21,8041,742]
d: [2,8043,2,9,3,5,21,8042,737]
d: [2,9773,2,7,2,6,21,8043,1717]
d: [2,4293,2,23,6,3,44,1981,2606]
d: [2,9793,2,19,5,4,21,4293,2222]
d: [2,9794,2,17,4,5,21,9793,742]
d: [2,4317,2,21,6,2,44,1991,906]
d: [2,4626,2,19,5,3,21,4317,742]
d: [2,4627,2,15,4,4,21,4626,1986]
d: [2,9796,2,11,3,5,21,4627,2049]
d: [2,9797,2,9,2,6,21,9796,1717]
d: [2,9634,3,12,6,2,44,2054,906]
d: [2,10123,3,10,5,3,21,9634,1172]
d: [2,10124,3,8,4,4,21,10123,747]
d: [2,10125,3,6,3,5,21,10124,752]
d: [2,10126,3,4,2,6,21,10125,737]
d: [2,9390,3,12,6,2,44,2049,906]
d: [2,10128,3,10,5,3,21,9390,1549]
d: [2,10129,3,8,4,4,21,10128,747]
d: [2,10130,3,6,3,5,21,10129,752]
d: [2,10131,3,4,2,6,21,10130,737]
d: [2,7971,2,17,6,2,44,3500,906]
d: [2,7980,2,14,5,2,19,7971]
d: [2,10152,2,12,4,3,21,7980,747]
d: [2,10153,2,8,3,4,21,10152,3081]
d: [2,10162,3,2,1,7,21,10033,1717]
d: [2,10164,3,6,3,2,21,10013,742]
d: [2,10165,3,4,2,3,21,10164,1695]
d: [2,10166,3,2,1,4,21,10165,752]
d: [2,9931,2,13,6,3,44,9892,2606]
d: [2,9936,2,10,5,3,19,9931]
d: [2,10175,2,8,4,4,21,9936,787]
d: [2,10176,2,6,3,5,21,10175,742]
d: [2,10177,2,4,2,6,21,10176,762]
d: [2,10178,2,2,1,7,21,10177,772]
d: [2,4291,2,23,6,2,44,1981,926]
d: [2,10203,2,19,5,3,21,4291,2222]
d: [2,10204,2,17,4,4,21,10203,742]
d: [2,10310,1,15,3,1,44,732,495]
d: [2,10336,1,13,2,2,21,10310,752]
d: [2,10337,1,11,1,3,21,10336,747]
d: [2,10322,1,17,3,1,44,732,495]
d: [2,10364,1,15,2,2,21,10322,742]
d: [2,10365,1,11,1,3,21,10364,1986]
d: [2,10328,1,23,6,2,44,732,926]
d: [2,10372,1,19,5,3,21,10328,1986]
d: [2,10373,1,17,4,4,21,10372,742]
d: [2,10374,1,13,3,5,21,10373,2049]
d: [2,10507,2,19,6,2,44,3545,922]
d: [2,10514,2,16,5,2,19,10507]
d: [2,10523,2,14,4,3,21,10514,747]
d: [2,10524,2,10,3,4,21,10523,3081]
d: [2,10525,2,8,2,5,21,10524,1705]
d: [2,10526,2,4,1,6,21,10525,3078]
d: [2,10640,3,2,1,7,21,10636,1722]
d: [2,10629,3,12,6,3,44,3093,2606]
d: [2,10660,3,10,5,4,21,10629,1190]
d: [2,10661,3,8,4,5,21,10660,752]
d: [2,10662,3,6,3,6,21,10661,742]
d: [2,10663,3,4,2,7,21,10662,737]
d: [2,10664,3,2,1,8,21,10663,1722]
d: [2,10565,3,8,4,1,40,2045,491]
d: [2,10675,3,6,3,2,21,10565,752]
d: [2,10676,3,4,2,3,21,10675,747]
d: [2,9851,1,15,6,2,44,7815,922]
d: [2,10714,1,10,5,3,21,9851,9891]
d: [2,10715,1,8,4,4,21,10714,747]
d: [2,10716,1,6,3,5,21,10715,752]
d: [2,10717,1,4,2,6,21,10716,737]
d: [2,10475,3,12,6,2,44,2042,906]
d: [2,10746,3,10,5,3,21,10475,1538]
d: [2,10747,3,8,4,4,21,10746,747]
d: [2,10748,3,6,3,5,21,10747,752]
d: [2,10749,3,4,2,6,21,10748,737]
d: [2,10470,3,8,4,1,40,2042,485]
d: [2,10751,3,6,3,2,21,10470,752]
d: [2,10752,3,4,2,3,21,10751,747]
f: 10752,10749,10717,10676,10664,10640,10526,10374,10365,10337,10204,10178,10166,10162,10153,10131,10126,9797,9794,9773,9668,9628,9513,9401,7919,4979,8064,8028,7774,7668,3686,7643,7623,7329,7279,7195,7114,7052,6940,6924,6901,6827,6785,6572,5325,4917,4793,4618,4356,3529,3503,3460,3438,3422,3262,3177,3106,3103,3019,2288,2200,2112,2036,2026,1801,1817,1780,1778,1522,1439,1464,1521,1407,1192,1176,1168,653,652,651,650,645,644,643,642,702,727,287,286,285,283,282,280,279,278,277,272,273,274,275,276
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_3__l92_xxreal_3
% SZS output start Proof for xxreal_3__l92_xxreal_3
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (~v1_xreal_0(X0) => k4_xxreal_3(X1,k4_xxreal_3(X0,X2)) = k4_xxreal_3(k4_xxreal_3(X1,X0),X2))))) [input l92_xxreal_3]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (~v1_xreal_0(X0) => k4_xxreal_3(X1,k4_xxreal_3(X0,X2)) = k4_xxreal_3(k4_xxreal_3(X1,X0),X2))))) [negated conjecture 1]
9. ! [X0] : ((~v1_xreal_0(X0) & ~v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (v2_xxreal_0(X0) & v1_xxreal_0(X0))) [input cc1_xxreal_3]
12. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
13. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
20. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) => (v2_xxreal_0(X0) & v1_xxreal_0(X0))) [input cc4_xxreal_0]
25. ! [X0] : (v1_xboole_0(X0) => v7_ordinal1(X0)) [input cc7_ordinal1]
26. ! [X0] : ((v1_xxreal_0(X0) & v1_xboole_0(X0)) => (~v3_xxreal_0(X0) & ~v2_xxreal_0(X0) & v1_xxreal_0(X0))) [input cc7_xxreal_0]
27. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
31. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => k4_xxreal_3(X1,X0) = k4_xxreal_3(X0,X1)) [input commutativity_k4_xxreal_3]
32. k1_numbers = k1_xxreal_0 [input d2_xxreal_0]
34. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (~(~(k4_xxreal_3(X0,X1) = X2 <=> k6_numbers = X2) & ~(((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ~(((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ~(~(k4_xxreal_3(X0,X1) = X2 <=> k2_xxreal_0 = X2) & ((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ~(~(k4_xxreal_3(X0,X1) = X2 <=> k1_xxreal_0 = X2) & ((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ((v1_xreal_0(X1) & v1_xreal_0(X0)) => (k4_xxreal_3(X0,X1) = X2 <=> ? [X3] : (? [X4] : (k3_xcmplx_0(X3,X4) = X2 & X1 = X4 & X0 = X3 & v1_xcmplx_0(X4)) & v1_xcmplx_0(X3)))))))) [input d5_xxreal_3]
43. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v1_xxreal_0(k4_xxreal_3(X0,X1))) [input dt_k4_xxreal_3]
50. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
51. ~v1_xreal_0(k2_xxreal_0) [input fc1_xreal_0]
52. v1_xxreal_0(k1_xxreal_0) [input fc1_xxreal_0]
55. ! [X0,X1] : ((v3_xxreal_0(X1) & v1_xxreal_0(X1) & v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc24_xxreal_3]
57. ! [X0,X1] : ((v3_xxreal_0(X1) & v1_xxreal_0(X1) & v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc25_xxreal_3]
59. ! [X0,X1] : ((v2_xxreal_0(X1) & v1_xxreal_0(X1) & v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc26_xxreal_3]
60. ! [X0,X1] : ((~v3_xxreal_0(X1) & v1_xxreal_0(X1) & ~v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc27_xxreal_3]
62. ! [X0,X1] : ((~v3_xxreal_0(X1) & v1_xxreal_0(X1) & ~v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc29_xxreal_3]
63. ~v1_xreal_0(k1_xxreal_0) [input fc2_xreal_0]
64. v1_xxreal_0(k2_xxreal_0) [input fc2_xxreal_0]
65. ! [X0,X1] : ((v1_xxreal_0(X1) & ~v1_xboole_0(X1) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) => (v1_xxreal_0(k4_xxreal_3(X0,X1)) & ~v1_xboole_0(k4_xxreal_3(X0,X1)))) [input fc36_xxreal_3]
69. v3_xxreal_0(k2_xxreal_0) [input fc4_xxreal_0]
73. ! [X0,X1,X2,X3] : ((v1_xcmplx_0(X3) & v1_xcmplx_0(X2) & v1_xreal_0(X1) & v1_xreal_0(X0)) => ((X1 = X3 & X0 = X2) => k4_xxreal_3(X0,X1) = k3_xcmplx_0(X2,X3))) [input ie4_xxreal_3]
95. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
100. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
102. ! [X0] : (v1_xcmplx_0(X0) => k6_numbers = k3_xcmplx_0(X0,k6_numbers)) [input t2_arithm]
120. ? [X0] : (? [X1] : (? [X2] : ((k4_xxreal_3(X1,k4_xxreal_3(X0,X2)) != k4_xxreal_3(k4_xxreal_3(X1,X0),X2) & ~v1_xreal_0(X0)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
121. ? [X0] : (? [X1] : (? [X2] : (k4_xxreal_3(X1,k4_xxreal_3(X0,X2)) != k4_xxreal_3(k4_xxreal_3(X1,X0),X2) & ~v1_xreal_0(X0) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 120]
125. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
126. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 13]
129. ! [X0] : (k6_numbers = k3_xcmplx_0(X0,k6_numbers) | ~v1_xcmplx_0(X0)) [ennf transformation 102]
130. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
133. ! [X0] : (v7_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 25]
138. ! [X0] : (! [X1] : (! [X2] : ((((k4_xxreal_3(X0,X1) = X2 <=> k6_numbers = X2) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> k2_xxreal_0 = X2) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> k1_xxreal_0 = X2) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> ? [X3] : (? [X4] : (k3_xcmplx_0(X3,X4) = X2 & X1 = X4 & X0 = X3 & v1_xcmplx_0(X4)) & v1_xcmplx_0(X3))) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 34]
139. ! [X0] : (! [X1] : (! [X2] : ((((k4_xxreal_3(X0,X1) = X2 <=> k6_numbers = X2) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> k2_xxreal_0 = X2) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> k1_xxreal_0 = X2) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> ? [X3] : (? [X4] : (k3_xcmplx_0(X3,X4) = X2 & X1 = X4 & X0 = X3 & v1_xcmplx_0(X4)) & v1_xcmplx_0(X3))) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 138]
142. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 27]
146. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | (v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0))) [ennf transformation 20]
147. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0)) [flattening 146]
152. ! [X0] : ((~v3_xxreal_0(X0) & ~v2_xxreal_0(X0) & v1_xxreal_0(X0)) | (~v1_xxreal_0(X0) | ~v1_xboole_0(X0))) [ennf transformation 26]
153. ! [X0] : ((~v3_xxreal_0(X0) & ~v2_xxreal_0(X0) & v1_xxreal_0(X0)) | ~v1_xxreal_0(X0) | ~v1_xboole_0(X0)) [flattening 152]
154. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | (v1_xreal_0(X0) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 9]
155. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | v1_xreal_0(X0) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 154]
182. ! [X0,X1] : ((v1_xxreal_0(k4_xxreal_3(X0,X1)) & ~v1_xboole_0(k4_xxreal_3(X0,X1))) | (~v1_xxreal_0(X1) | v1_xboole_0(X1) | ~v1_xxreal_0(X0) | v1_xboole_0(X0))) [ennf transformation 65]
183. ! [X0,X1] : ((v1_xxreal_0(k4_xxreal_3(X0,X1)) & ~v1_xboole_0(k4_xxreal_3(X0,X1))) | ~v1_xxreal_0(X1) | v1_xboole_0(X1) | ~v1_xxreal_0(X0) | v1_xboole_0(X0)) [flattening 182]
200. ! [X0,X1] : ((~v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 62]
201. ! [X0,X1] : ((~v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 200]
202. ! [X0,X1] : ((~v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 60]
203. ! [X0,X1] : ((~v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 202]
206. ! [X0,X1] : ((v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 57]
207. ! [X0,X1] : ((v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | ~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 206]
208. ! [X0,X1] : ((v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 55]
209. ! [X0,X1] : ((v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | ~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 208]
210. ! [X0,X1] : ((v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (~v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 59]
211. ! [X0,X1] : ((v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | ~v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 210]
212. ! [X0,X1] : (v1_xxreal_0(k4_xxreal_3(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 43]
213. ! [X0,X1] : (v1_xxreal_0(k4_xxreal_3(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 212]
214. ! [X0,X1] : (k4_xxreal_3(X1,X0) = k4_xxreal_3(X0,X1) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 31]
215. ! [X0,X1] : (k4_xxreal_3(X1,X0) = k4_xxreal_3(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 214]
218. ! [X0,X1,X2,X3] : ((k4_xxreal_3(X0,X1) = k3_xcmplx_0(X2,X3) | (X1 != X3 | X0 != X2)) | (~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 73]
219. ! [X0,X1,X2,X3] : (k4_xxreal_3(X0,X1) = k3_xcmplx_0(X2,X3) | X1 != X3 | X0 != X2 | ~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 218]
220. ? [X0] : (? [X1] : (? [X2] : (k4_xxreal_3(X1,k4_xxreal_3(X0,X2)) != k4_xxreal_3(k4_xxreal_3(X1,X0),X2) & ~v1_xreal_0(X0) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (k4_xxreal_3(X1,k4_xxreal_3(sK0,X2)) != k4_xxreal_3(k4_xxreal_3(X1,sK0),X2) & ~v1_xreal_0(sK0) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
221. ? [X1] : (? [X2] : (k4_xxreal_3(X1,k4_xxreal_3(sK0,X2)) != k4_xxreal_3(k4_xxreal_3(X1,sK0),X2) & ~v1_xreal_0(sK0) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (k4_xxreal_3(sK1,k4_xxreal_3(sK0,X2)) != k4_xxreal_3(k4_xxreal_3(sK1,sK0),X2) & ~v1_xreal_0(sK0) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
222. ? [X2] : (k4_xxreal_3(sK1,k4_xxreal_3(sK0,X2)) != k4_xxreal_3(k4_xxreal_3(sK1,sK0),X2) & ~v1_xreal_0(sK0) & v1_xxreal_0(X2)) => (k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(k4_xxreal_3(sK1,sK0),sK2) & ~v1_xreal_0(sK0) & v1_xxreal_0(sK2)) [choice axiom]
223. ((k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(k4_xxreal_3(sK1,sK0),sK2) & ~v1_xreal_0(sK0) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 121,222,221,220]
224. ! [X0] : (! [X1] : (! [X2] : (((((k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2) & (k6_numbers = X2 | k4_xxreal_3(X0,X1) != X2)) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2) & (k2_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2) & (k1_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | ! [X3] : (! [X4] : (k3_xcmplx_0(X3,X4) != X2 | X1 != X4 | X0 != X3 | ~v1_xcmplx_0(X4)) | ~v1_xcmplx_0(X3))) & (? [X3] : (? [X4] : (k3_xcmplx_0(X3,X4) = X2 & X1 = X4 & X0 = X3 & v1_xcmplx_0(X4)) & v1_xcmplx_0(X3)) | k4_xxreal_3(X0,X1) != X2)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 139]
225. ! [X0] : (! [X1] : (! [X2] : (((((k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2) & (k6_numbers = X2 | k4_xxreal_3(X0,X1) != X2)) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2) & (k2_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2) & (k1_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | ! [X3] : (! [X4] : (k3_xcmplx_0(X3,X4) != X2 | X1 != X4 | X0 != X3 | ~v1_xcmplx_0(X4)) | ~v1_xcmplx_0(X3))) & (? [X5] : (? [X6] : (k3_xcmplx_0(X5,X6) = X2 & X1 = X6 & X0 = X5 & v1_xcmplx_0(X6)) & v1_xcmplx_0(X5)) | k4_xxreal_3(X0,X1) != X2)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [rectify 224]
226. ! [X2,X1,X0] : (? [X5] : (? [X6] : (k3_xcmplx_0(X5,X6) = X2 & X1 = X6 & X0 = X5 & v1_xcmplx_0(X6)) & v1_xcmplx_0(X5)) => (? [X6] : (k3_xcmplx_0(sK3(X0,X1,X2),X6) = X2 & X1 = X6 & sK3(X0,X1,X2) = X0 & v1_xcmplx_0(X6)) & v1_xcmplx_0(sK3(X0,X1,X2)))) [choice axiom]
227. ! [X2,X1,X0] : (? [X6] : (k3_xcmplx_0(sK3(X0,X1,X2),X6) = X2 & X1 = X6 & sK3(X0,X1,X2) = X0 & v1_xcmplx_0(X6)) => (k3_xcmplx_0(sK3(X0,X1,X2),sK4(X0,X1,X2)) = X2 & sK4(X0,X1,X2) = X1 & sK3(X0,X1,X2) = X0 & v1_xcmplx_0(sK4(X0,X1,X2)))) [choice axiom]
228. ! [X0] : (! [X1] : (! [X2] : (((((k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2) & (k6_numbers = X2 | k4_xxreal_3(X0,X1) != X2)) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2) & (k2_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2) & (k1_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | ! [X3] : (! [X4] : (k3_xcmplx_0(X3,X4) != X2 | X1 != X4 | X0 != X3 | ~v1_xcmplx_0(X4)) | ~v1_xcmplx_0(X3))) & (((k3_xcmplx_0(sK3(X0,X1,X2),sK4(X0,X1,X2)) = X2 & sK4(X0,X1,X2) = X1 & sK3(X0,X1,X2) = X0 & v1_xcmplx_0(sK4(X0,X1,X2))) & v1_xcmplx_0(sK3(X0,X1,X2))) | k4_xxreal_3(X0,X1) != X2)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [skolemisation 225,227,226]
272. v1_xxreal_0(sK0) [cnf transformation 223] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
273. v1_xxreal_0(sK1) [cnf transformation 223] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
274. v1_xxreal_0(sK2) [cnf transformation 223] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
275. ~v1_xreal_0(sK0) [cnf transformation 223] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
276. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(k4_xxreal_3(sK1,sK0),sK2) [cnf transformation 223] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
277. ~v1_xreal_0(k1_xxreal_0) [cnf transformation 63] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
278. ~v1_xreal_0(k2_xxreal_0) [cnf transformation 51] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
279. v1_xboole_0(k1_xboole_0) [cnf transformation 50] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
280. v3_xxreal_0(k2_xxreal_0) [cnf transformation 69] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
282. v1_xxreal_0(k1_xxreal_0) [cnf transformation 52] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
283. v1_xxreal_0(k2_xxreal_0) [cnf transformation 64] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
285. k6_numbers = k1_xboole_0 [cnf transformation 95] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
286. k1_numbers = k1_xxreal_0 [cnf transformation 32] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
287. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 100] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
296. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 125] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
297. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 126] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
301. k6_numbers = k3_xcmplx_0(X0,k6_numbers) | ~v1_xcmplx_0(X0) [cnf transformation 129] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
302. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 130] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
305. ~v1_xboole_0(X0) | v7_ordinal1(X0) [cnf transformation 133] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
319. k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2 | ~v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
320. k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2 | ~v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
321. k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2 | ~v3_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
322. k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2 | ~v3_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
327. k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2 | ~v3_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
328. k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2 | ~v3_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
329. k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2 | ~v2_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
330. k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2 | ~v2_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
394. k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2 | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:24,wCS:9600,sine:255,thAx:0,allAx:3,thDist:-3}
429. k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2 | v2_xxreal_0(X1) | v3_xxreal_0(X1) | v3_xxreal_0(X1) | v2_xxreal_0(X1) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 228] {a:0,w:24,wCS:9600,sine:255,thAx:0,allAx:3,thDist:-3}
432. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 142] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
437. ~v1_xxreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xboole_0(X0) [cnf transformation 147] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
442. ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) | ~v2_xxreal_0(X0) [cnf transformation 153] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
443. ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) | ~v3_xxreal_0(X0) [cnf transformation 153] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
445. ~v1_xxreal_0(X0) | v1_xreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 155] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
471. ~v1_xboole_0(k4_xxreal_3(X0,X1)) | ~v1_xxreal_0(X1) | v1_xboole_0(X1) | ~v1_xxreal_0(X0) | v1_xboole_0(X0) [cnf transformation 183] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
483. ~v3_xxreal_0(k4_xxreal_3(X0,X1)) | v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 201] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
485. ~v2_xxreal_0(k4_xxreal_3(X0,X1)) | v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 203] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
489. v2_xxreal_0(k4_xxreal_3(X0,X1)) | ~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 207] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
491. v3_xxreal_0(k4_xxreal_3(X0,X1)) | ~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 209] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
493. v2_xxreal_0(k4_xxreal_3(X0,X1)) | ~v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 211] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
494. v1_xxreal_0(k4_xxreal_3(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 213] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
495. k4_xxreal_3(X1,X0) = k4_xxreal_3(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 215] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
497. k4_xxreal_3(X0,X1) = k3_xcmplx_0(X2,X3) | X1 != X3 | X0 != X2 | ~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 219] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:1,thDist:-1}
543. k4_xxreal_3(X0,X1) = k6_numbers | v2_xxreal_0(X1) | v3_xxreal_0(X1) | v3_xxreal_0(X1) | v2_xxreal_0(X1) | v1_xreal_0(X0) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 429] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:3,thDist:-3}
578. k4_xxreal_3(X0,X1) = k6_numbers | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 394] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:3,thDist:-3}
642. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v2_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 330] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
643. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v2_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 329] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
644. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v3_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 328] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
645. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v3_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 327] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
650. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v3_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(k1_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 322] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
651. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v3_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(k1_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 321] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
652. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(k1_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 320] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
653. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(k1_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 319] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
667. k3_xcmplx_0(X2,X3) = k4_xxreal_3(X0,X3) | X0 != X2 | ~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X3) | ~v1_xreal_0(X0) [equality resolution 497] {a:0,w:18,wCS:7200,sine:255,thAx:0,allAx:1,thDist:-1}
668. k3_xcmplx_0(X2,X3) = k4_xxreal_3(X2,X3) | ~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) [equality resolution 667] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:1,thDist:-1}
702. k4_xxreal_3(X0,X1) = k6_numbers | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [duplicate literal removal 578] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
727. k4_xxreal_3(X0,X1) = k6_numbers | v2_xxreal_0(X1) | v3_xxreal_0(X1) | v1_xreal_0(X0) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [duplicate literal removal 543] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
730. 1 <=> k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) = k4_xxreal_3(k4_xxreal_3(sK1,sK0),sK2) [avatar definition]
732. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(k4_xxreal_3(sK1,sK0),sK2) <- (~1) [avatar component clause 730] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
733. ~1 [avatar split clause 276,730]
735. 2 <=> v1_xreal_0(sK0) [avatar definition]
737. ~v1_xreal_0(sK0) <- (~2) [avatar component clause 735] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
738. ~2 [avatar split clause 275,735]
740. 3 <=> v1_xxreal_0(sK2) [avatar definition]
742. v1_xxreal_0(sK2) <- (3) [avatar component clause 740] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
743. 3 [avatar split clause 274,740]
745. 4 <=> v1_xxreal_0(sK1) [avatar definition]
747. v1_xxreal_0(sK1) <- (4) [avatar component clause 745] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
748. 4 [avatar split clause 273,745]
750. 5 <=> v1_xxreal_0(sK0) [avatar definition]
752. v1_xxreal_0(sK0) <- (5) [avatar component clause 750] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
753. 5 [avatar split clause 272,750]
755. 6 <=> v1_xreal_0(k1_xxreal_0) [avatar definition]
757. ~v1_xreal_0(k1_xxreal_0) <- (~6) [avatar component clause 755] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
758. ~6 [avatar split clause 277,755]
760. 7 <=> v1_xreal_0(k2_xxreal_0) [avatar definition]
762. ~v1_xreal_0(k2_xxreal_0) <- (~7) [avatar component clause 760] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
763. ~7 [avatar split clause 278,760]
765. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
767. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 765] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
768. 8 [avatar split clause 279,765]
770. 9 <=> v3_xxreal_0(k2_xxreal_0) [avatar definition]
772. v3_xxreal_0(k2_xxreal_0) <- (9) [avatar component clause 770] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
773. 9 [avatar split clause 280,770]
780. 11 <=> v1_xxreal_0(k1_xxreal_0) [avatar definition]
783. 11 [avatar split clause 282,780]
785. 12 <=> v1_xxreal_0(k2_xxreal_0) [avatar definition]
787. v1_xxreal_0(k2_xxreal_0) <- (12) [avatar component clause 785] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
788. 12 [avatar split clause 283,785]
795. 14 <=> k6_numbers = k1_xboole_0 [avatar definition]
797. k6_numbers = k1_xboole_0 <- (14) [avatar component clause 795] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
798. 14 [avatar split clause 285,795]
800. 15 <=> k1_numbers = k1_xxreal_0 [avatar definition]
802. k1_numbers = k1_xxreal_0 <- (15) [avatar component clause 800] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
803. 15 [avatar split clause 286,800]
805. 16 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
807. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (16) [avatar component clause 805] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
808. 16 [avatar split clause 287,805]
835. 22 <=> v1_xxreal_0(k6_numbers) [avatar definition]
843. 24 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k6_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | v3_xxreal_0(X1) | v2_xxreal_0(X1)) [avatar definition]
844. k4_xxreal_3(X0,X1) = k6_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | v3_xxreal_0(X1) | v2_xxreal_0(X1) <- (24) [avatar component clause 843] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
845. ~22 | 24 [avatar split clause 727,843,835]
905. 35 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k6_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | v2_xxreal_0(X0) | v3_xxreal_0(X0)) [avatar definition]
906. k4_xxreal_3(X0,X1) = k6_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | v2_xxreal_0(X0) | v3_xxreal_0(X0) <- (35) [avatar component clause 905] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
907. ~22 | 35 [avatar split clause 702,905,835]
913. 37 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v3_xxreal_0(X0) | ~v2_xxreal_0(X1)) [avatar definition]
914. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v3_xxreal_0(X0) | ~v2_xxreal_0(X1) <- (37) [avatar component clause 913] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
915. ~12 | 37 [avatar split clause 642,913,785]
917. 38 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v2_xxreal_0(X1)) [avatar definition]
918. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v2_xxreal_0(X1) <- (38) [avatar component clause 917] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
919. ~12 | 38 [avatar split clause 643,917,785]
921. 39 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v2_xxreal_0(X0) | ~v3_xxreal_0(X1)) [avatar definition]
922. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v2_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (39) [avatar component clause 921] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
923. ~12 | 39 [avatar split clause 644,921,785]
925. 40 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v3_xxreal_0(X1)) [avatar definition]
926. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (40) [avatar component clause 925] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
927. ~12 | 40 [avatar split clause 645,925,785]
929. 41 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v3_xxreal_0(X0) | ~v3_xxreal_0(X1)) [avatar definition]
930. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v3_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (41) [avatar component clause 929] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:3,thDist:-3}
931. ~11 | 41 [avatar split clause 650,929,780]
933. 42 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v3_xxreal_0(X1)) [avatar definition]
934. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (42) [avatar component clause 933] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:3,thDist:-3}
935. ~11 | 42 [avatar split clause 651,933,780]
937. 43 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v2_xxreal_0(X0) | ~v2_xxreal_0(X1)) [avatar definition]
938. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v2_xxreal_0(X0) | ~v2_xxreal_0(X1) <- (43) [avatar component clause 937] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:3,thDist:-3}
939. ~11 | 43 [avatar split clause 652,937,780]
941. 44 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v2_xxreal_0(X1)) [avatar definition]
942. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v2_xxreal_0(X1) <- (44) [avatar component clause 941] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:3,thDist:-3}
943. ~11 | 44 [avatar split clause 653,941,780]
1168. v1_xboole_0(k6_numbers) <- (8, 14) [backward demodulation 767,797] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
1170. 90 <=> v1_xboole_0(k6_numbers) [avatar definition]
1172. v1_xboole_0(k6_numbers) <- (90) [avatar component clause 1170] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1173. 90 | ~8 | ~14 [avatar split clause 1168,795,765,1170]
1176. ~v1_xreal_0(k1_numbers) <- (~6, 15) [backward demodulation 757,802] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
1188. 93 <=> v1_xreal_0(k1_numbers) [avatar definition]
1190. ~v1_xreal_0(k1_numbers) <- (~93) [avatar component clause 1188] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1191. ~93 | 6 | ~15 [avatar split clause 1176,800,755,1188]
1192. m1_subset_1(k6_numbers,k4_ordinal1) <- (14, 16) [forward demodulation 807,797] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
1194. 94 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
1196. m1_subset_1(k6_numbers,k4_ordinal1) <- (94) [avatar component clause 1194] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1197. 94 | ~14 | ~16 [avatar split clause 1192,805,795,1194]
1351. 117 <=> v7_ordinal1(k6_numbers) [avatar definition]
1353. v7_ordinal1(k6_numbers) <- (117) [avatar component clause 1351] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1407. v7_ordinal1(k6_numbers) <- (94) [resolution 432,1196] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1409. 117 | ~94 [avatar split clause 1407,1194,1351]
1439. ~v1_xxreal_0(k6_numbers) | ~v2_xxreal_0(k6_numbers) <- (90) [resolution 442,1172] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:2,thDist:-2}
1464. ~v1_xxreal_0(k6_numbers) | ~v3_xxreal_0(k6_numbers) <- (90) [resolution 443,1172] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:2,thDist:-2}
1521. v1_xxreal_0(k6_numbers) <- (117) [resolution 1353,297] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1522. v1_xreal_0(k6_numbers) <- (117) [resolution 1353,296] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1529. 22 | ~117 [avatar split clause 1521,1351,835]
1531. 142 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1533. ~v3_xxreal_0(k6_numbers) <- (~142) [avatar component clause 1531] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1534. ~142 | ~22 | ~90 [avatar split clause 1464,1170,835,1531]
1536. 143 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1538. ~v2_xxreal_0(k6_numbers) <- (~143) [avatar component clause 1536] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1539. ~143 | ~22 | ~90 [avatar split clause 1439,1170,835,1536]
1547. 145 <=> v1_xreal_0(k6_numbers) [avatar definition]
1549. v1_xreal_0(k6_numbers) <- (145) [avatar component clause 1547] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
1550. 145 | ~117 [avatar split clause 1522,1351,1547]
1686. 162 <=> v1_xboole_0(sK0) [avatar definition]
1687. ~v1_xboole_0(sK0) <- (~162) [avatar component clause 1686] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1688. v1_xboole_0(sK0) <- (162) [avatar component clause 1686] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1690. 163 <=> v2_xxreal_0(sK0) [avatar definition]
1694. 164 <=> v3_xxreal_0(sK0) [avatar definition]
1695. ~v3_xxreal_0(sK0) <- (~164) [avatar component clause 1694] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1699. 165 <=> v1_xboole_0(sK1) [avatar definition]
1700. ~v1_xboole_0(sK1) <- (~165) [avatar component clause 1699] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1701. v1_xboole_0(sK1) <- (165) [avatar component clause 1699] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1703. 166 <=> v2_xxreal_0(sK1) [avatar definition]
1705. v2_xxreal_0(sK1) <- (166) [avatar component clause 1703] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1707. 167 <=> v3_xxreal_0(sK1) [avatar definition]
1708. ~v3_xxreal_0(sK1) <- (~167) [avatar component clause 1707] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1709. v3_xxreal_0(sK1) <- (167) [avatar component clause 1707] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1716. 169 <=> v2_xxreal_0(sK2) [avatar definition]
1717. ~v2_xxreal_0(sK2) <- (~169) [avatar component clause 1716] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1718. v2_xxreal_0(sK2) <- (169) [avatar component clause 1716] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1720. 170 <=> v3_xxreal_0(sK2) [avatar definition]
1721. ~v3_xxreal_0(sK2) <- (~170) [avatar component clause 1720] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1722. v3_xxreal_0(sK2) <- (170) [avatar component clause 1720] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1773. ~v1_xxreal_0(sK0) | ~v3_xxreal_0(sK0) <- (162) [resolution 1688,443] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1774. ~v1_xxreal_0(sK0) | ~v2_xxreal_0(sK0) <- (162) [resolution 1688,442] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1778. ~v3_xxreal_0(sK0) <- (5, 162) [subsumption resolution 1773,752] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1779. ~164 | ~5 | ~162 [avatar split clause 1778,1686,750,1694]
1780. ~v2_xxreal_0(sK0) <- (5, 162) [subsumption resolution 1774,752] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1781. ~163 | ~5 | ~162 [avatar split clause 1780,1686,750,1690]
1800. v1_xreal_0(sK0) | v3_xxreal_0(sK0) | v2_xxreal_0(sK0) <- (5) [resolution 445,752] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1801. v1_xreal_0(sK1) | v3_xxreal_0(sK1) | v2_xxreal_0(sK1) <- (4) [resolution 445,747] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1817. v3_xxreal_0(sK0) | v2_xxreal_0(sK0) <- (~2, 5) [subsumption resolution 1800,737] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1818. 163 | 164 | 2 | ~5 [avatar split clause 1817,750,735,1694,1690]
1820. 185 <=> v1_xreal_0(sK1) [avatar definition]
1821. ~v1_xreal_0(sK1) <- (~185) [avatar component clause 1820] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1822. v1_xreal_0(sK1) <- (185) [avatar component clause 1820] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1823. 166 | 167 | 185 | ~4 [avatar split clause 1801,745,1820,1707,1703]
1825. 186 <=> v1_xreal_0(sK2) [avatar definition]
1827. v1_xreal_0(sK2) <- (186) [avatar component clause 1825] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1887. 194 <=> v7_ordinal1(sK1) [avatar definition]
1889. v7_ordinal1(sK1) <- (194) [avatar component clause 1887] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1979. 201 <=> k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) = k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) [avatar definition]
1981. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) <- (~201) [avatar component clause 1979] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1985. 202 <=> v1_xxreal_0(k4_xxreal_3(sK1,sK0)) [avatar definition]
1986. v1_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (202) [avatar component clause 1985] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1987. ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (~202) [avatar component clause 1985] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1989. 203 <=> k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) = k4_xxreal_3(sK2,k4_xxreal_3(sK1,sK0)) [avatar definition]
1991. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(sK2,k4_xxreal_3(sK1,sK0)) <- (~203) [avatar component clause 1989] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2021. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~202) [resolution 1987,494] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2025. ~v1_xxreal_0(sK1) <- (5, ~202) [subsumption resolution 2021,752] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2026. $false <- (4, 5, ~202) [subsumption resolution 2025,747] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2027. ~4 | ~5 | 202 [avatar contradiction clause 2026]
2036. v3_xxreal_0(k4_xxreal_3(sK1,sK0)) | v2_xxreal_0(k4_xxreal_3(sK1,sK0)) | v1_xboole_0(k4_xxreal_3(sK1,sK0)) <- (202) [resolution 1986,437] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2040. 205 <=> v2_xxreal_0(k4_xxreal_3(sK1,sK0)) [avatar definition]
2041. ~v2_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (~205) [avatar component clause 2040] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2042. v2_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (205) [avatar component clause 2040] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2044. 206 <=> v3_xxreal_0(k4_xxreal_3(sK1,sK0)) [avatar definition]
2045. ~v3_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (~206) [avatar component clause 2044] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2046. v3_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (206) [avatar component clause 2044] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2048. 207 <=> v1_xreal_0(k4_xxreal_3(sK1,sK0)) [avatar definition]
2049. ~v1_xreal_0(k4_xxreal_3(sK1,sK0)) <- (~207) [avatar component clause 2048] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2050. v1_xreal_0(k4_xxreal_3(sK1,sK0)) <- (207) [avatar component clause 2048] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2053. 208 <=> v1_xboole_0(k4_xxreal_3(sK1,sK0)) [avatar definition]
2054. ~v1_xboole_0(k4_xxreal_3(sK1,sK0)) <- (~208) [avatar component clause 2053] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2055. v1_xboole_0(k4_xxreal_3(sK1,sK0)) <- (208) [avatar component clause 2053] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2056. 208 | 205 | 206 | ~202 [avatar split clause 2036,1985,2044,2040,2053]
2080. 210 <=> v2_xxreal_0(k4_xxreal_3(sK0,sK1)) [avatar definition]
2081. ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (~210) [avatar component clause 2080] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2105. v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) | v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (206) [resolution 2046,483] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2109. ~v1_xxreal_0(sK0) | v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (~164, 206) [subsumption resolution 2105,1695] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2110. v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (5, ~164, 206) [subsumption resolution 2109,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2111. ~v1_xxreal_0(sK1) <- (5, ~164, ~167, 206) [subsumption resolution 2110,1708] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2112. $false <- (4, 5, ~164, ~167, 206) [subsumption resolution 2111,747] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2113. ~4 | ~5 | 164 | 167 | ~206 [avatar contradiction clause 2112]
2130. 212 <=> v1_xreal_0(k4_xxreal_3(sK0,sK1)) [avatar definition]
2162. 213 <=> v3_xxreal_0(k4_xxreal_3(sK0,sK1)) [avatar definition]
2164. ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (~213) [avatar component clause 2162] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2187. ~v1_xxreal_0(sK0) | v1_xboole_0(sK0) | ~v1_xxreal_0(sK1) | v1_xboole_0(sK1) <- (208) [resolution 2055,471] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2197. v1_xboole_0(sK0) | ~v1_xxreal_0(sK1) | v1_xboole_0(sK1) <- (5, 208) [subsumption resolution 2187,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2198. ~v1_xxreal_0(sK1) | v1_xboole_0(sK1) <- (5, ~162, 208) [subsumption resolution 2197,1687] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2199. v1_xboole_0(sK1) <- (4, 5, ~162, 208) [subsumption resolution 2198,747] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2200. $false <- (4, 5, ~162, ~165, 208) [subsumption resolution 2199,1700] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2201. ~4 | ~5 | 162 | 165 | ~208 [avatar contradiction clause 2200]
2221. 214 <=> v1_xxreal_0(k4_xxreal_3(sK0,sK1)) [avatar definition]
2222. v1_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (214) [avatar component clause 2221] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
2223. ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (~214) [avatar component clause 2221] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
2243. ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (~213) [resolution 2164,491] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
2255. k3_xcmplx_0(X2,X3) = k4_xxreal_3(X2,X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) [subsumption resolution 668,302] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:1,thDist:-1}
2256. k3_xcmplx_0(X2,X3) = k4_xxreal_3(X2,X3) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) [subsumption resolution 2255,302] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
2264. k6_numbers = k4_xxreal_3(X0,k6_numbers) | ~v1_xcmplx_0(X0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X0) [superposition 301,2256] {a:1,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
2281. 217 <=> ! [X0] : (k6_numbers = k4_xxreal_3(X0,k6_numbers) | ~v1_xreal_0(X0)) [avatar definition]
2282. k6_numbers = k4_xxreal_3(X0,k6_numbers) | ~v1_xreal_0(X0) <- (217) [avatar component clause 2281] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2288. k6_numbers = k4_xxreal_3(X0,k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X0) [subsumption resolution 2264,302] {a:1,w:9,wCS:3600,sine:255,thAx:0,allAx:2,thDist:-2}
2289. ~145 | 217 [avatar split clause 2288,2281,1547]
2315. ~v1_xreal_0(k6_numbers) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK1) | v3_xxreal_0(sK0) | v2_xxreal_0(sK0) <- (24, ~207) [superposition 2049,844] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
2476. ~v3_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK1) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (38, ~206) [superposition 2045,918] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
2605. k1_numbers = k4_xxreal_3(X0,X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v3_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (15, 41) [forward demodulation 930,802] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
2606. k1_numbers = k4_xxreal_3(X0,X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (15, 42) [forward demodulation 934,802] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
2607. k1_numbers = k4_xxreal_3(X0,X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v2_xxreal_0(X0) | ~v2_xxreal_0(X1) <- (15, 43) [forward demodulation 938,802] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
2608. k1_numbers = k4_xxreal_3(X0,X1) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v2_xxreal_0(X1) <- (15, 44) [forward demodulation 942,802] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
2928. 235 <=> k1_numbers = k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) [avatar definition]
2929. k1_numbers = k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) <- (235) [avatar component clause 2928] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2930. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) <- (~235) [avatar component clause 2928] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2973. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) <- (~1, 15, 44) [superposition 732,2608] {a:1,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3015. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) <- (~1, 15, 44, 202) [subsumption resolution 2973,1986] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3016. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) <- (~1, 3, 15, 44, 202) [subsumption resolution 3015,742] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3017. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) <- (~1, 3, 15, 44, 202, ~207) [subsumption resolution 3016,2049] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3018. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(sK2) <- (~1, 3, 15, 44, 202, 205, ~207) [subsumption resolution 3017,2042] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3019. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) <- (~1, 3, 15, 44, 169, 202, 205, ~207) [subsumption resolution 3018,1718] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3020. ~235 | 1 | ~3 | ~15 | ~44 | ~169 | ~202 | ~205 | 207 [avatar split clause 3019,2048,2040,1985,1716,941,800,740,730,2928]
3049. k1_numbers != k1_numbers | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (15, 43, ~235) [superposition 2930,2607] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3050. k1_numbers != k1_numbers | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (15, 44, ~235) [superposition 2930,2608] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3051. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (15, 44, ~235) [trivial inequality removal 3050] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3052. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (15, 43, ~235) [trivial inequality removal 3049] {a:2,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3072. 240 <=> v2_xxreal_0(k4_xxreal_3(sK0,sK2)) [avatar definition]
3073. ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (~240) [avatar component clause 3072] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3074. v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (240) [avatar component clause 3072] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3076. 241 <=> v3_xxreal_0(k4_xxreal_3(sK0,sK2)) [avatar definition]
3077. ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (~241) [avatar component clause 3076] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3078. v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (241) [avatar component clause 3076] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3080. 242 <=> v1_xxreal_0(k4_xxreal_3(sK0,sK2)) [avatar definition]
3081. v1_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (242) [avatar component clause 3080] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3082. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (~242) [avatar component clause 3080] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3091. 244 <=> v1_xreal_0(k4_xxreal_3(sK0,sK2)) [avatar definition]
3093. v1_xreal_0(k4_xxreal_3(sK0,sK2)) <- (244) [avatar component clause 3091] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3102. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 43, ~235) [subsumption resolution 3052,747] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3103. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 43, 166, ~235) [subsumption resolution 3102,1705] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3104. ~240 | 244 | ~242 | ~4 | ~15 | ~43 | ~166 | 235 [avatar split clause 3103,2928,1703,937,800,745,3080,3091,3072]
3105. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 44, ~235) [subsumption resolution 3051,747] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3106. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 44, 166, ~235) [subsumption resolution 3105,1705] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3107. ~240 | 185 | ~242 | ~4 | ~15 | ~44 | ~166 | 235 [avatar split clause 3106,2928,1703,941,800,745,3080,1820,3072]
3164. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (~242) [resolution 3082,494] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3176. ~v1_xxreal_0(sK0) <- (3, ~242) [subsumption resolution 3164,742] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3177. $false <- (3, 5, ~242) [subsumption resolution 3176,752] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3178. ~3 | ~5 | 242 [avatar contradiction clause 3177]
3194. ~v3_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (~240) [resolution 3073,489] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3261. ~v1_xreal_0(k6_numbers) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK1) | v3_xxreal_0(sK0) | v2_xxreal_0(sK0) <- (4, 24, ~207) [subsumption resolution 2315,747] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
3262. ~v1_xreal_0(k6_numbers) | v1_xreal_0(sK1) | v3_xxreal_0(sK0) | v2_xxreal_0(sK0) <- (4, 5, 24, ~207) [subsumption resolution 3261,752] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
3263. 163 | 164 | 185 | ~145 | ~4 | ~5 | ~24 | 207 [avatar split clause 3262,2048,843,750,745,1547,1820,1694,1690]
3420. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK1) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (9, 38, ~206) [subsumption resolution 2476,772] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
3421. ~v1_xxreal_0(sK0) | v1_xreal_0(sK1) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (4, 9, 38, ~206) [subsumption resolution 3420,747] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
3422. v1_xreal_0(sK1) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (4, 5, 9, 38, ~206) [subsumption resolution 3421,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
3423. ~163 | ~167 | 185 | ~4 | ~5 | ~9 | ~38 | 206 [avatar split clause 3422,2044,917,770,750,745,1820,1707,1690]
3437. ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (4, ~213) [subsumption resolution 2243,747] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
3438. ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (4, 5, ~213) [subsumption resolution 3437,752] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
3439. ~163 | ~167 | ~4 | ~5 | 213 [avatar split clause 3438,2162,750,745,1707,1690]
3440. ~v3_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (~241) [resolution 3077,491] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3444. ~v3_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (38, ~241) [superposition 3077,918] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
3457. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (9, 38, ~241) [subsumption resolution 3444,772] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
3458. ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (5, 9, 38, ~241) [subsumption resolution 3457,752] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
3459. v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (3, 5, 9, 38, ~241) [subsumption resolution 3458,742] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
3460. ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (~2, 3, 5, 9, 38, ~241) [subsumption resolution 3459,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
3473. ~169 | ~164 | 2 | ~3 | ~5 | ~9 | ~38 | 241 [avatar split clause 3460,3076,917,770,750,740,735,1694,1716]
3498. 252 <=> k6_numbers = k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) [avatar definition]
3499. k6_numbers = k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) <- (252) [avatar component clause 3498] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3500. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) <- (~252) [avatar component clause 3498] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3502. ~v3_xxreal_0(sK2) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (3, ~241) [subsumption resolution 3440,742] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3503. ~v3_xxreal_0(sK2) | ~v2_xxreal_0(sK0) <- (3, 5, ~241) [subsumption resolution 3502,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3504. ~163 | ~170 | ~3 | ~5 | 241 [avatar split clause 3503,3076,750,740,1720,1690]
3528. ~v3_xxreal_0(sK2) | ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (3, ~240) [subsumption resolution 3194,742] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3529. ~v3_xxreal_0(sK2) | ~v3_xxreal_0(sK0) <- (3, 5, ~240) [subsumption resolution 3528,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3530. ~164 | ~170 | ~3 | ~5 | 240 [avatar split clause 3529,3072,750,740,1720,1694]
3543. 253 <=> k2_xxreal_0 = k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) [avatar definition]
3545. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) <- (~253) [avatar component clause 3543] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3686. v1_xreal_0(k4_xxreal_3(sK0,sK2)) | v3_xxreal_0(k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (242) [resolution 3081,445] {a:3,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:11,thDist:-11}
3719. v1_xreal_0(k1_numbers) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (15, 44, 244) [superposition 3093,2608] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
4035. 266 <=> k6_numbers = k4_xxreal_3(sK1,k6_numbers) [avatar definition]
4037. k6_numbers != k4_xxreal_3(sK1,k6_numbers) <- (~266) [avatar component clause 4035] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4289. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (38, ~201) [superposition 1981,918] {a:2,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4291. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (40, ~201) [superposition 1981,926] {a:2,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4293. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (15, 42, ~201) [superposition 1981,2606] {a:2,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4307. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(sK2,k6_numbers) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (35, ~203) [superposition 1991,906] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4316. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | v1_xreal_0(sK2) | v3_xxreal_0(k4_xxreal_3(sK1,sK0)) | v2_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (24, ~203) [superposition 1991,844] {a:2,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4317. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | v2_xxreal_0(sK2) | v3_xxreal_0(sK2) <- (35, ~203) [superposition 1991,906] {a:2,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4320. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (39, ~203) [superposition 1991,922] {a:2,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4343. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~214) [resolution 2223,494] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
4355. ~v1_xxreal_0(sK0) <- (4, ~214) [subsumption resolution 4343,747] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
4356. $false <- (4, 5, ~214) [subsumption resolution 4355,752] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
4357. ~4 | ~5 | 214 [avatar contradiction clause 4356]
4583. k6_numbers != k6_numbers | ~v1_xreal_0(sK1) <- (217, ~266) [superposition 4037,2282] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:12,thDist:-12}
4588. ~v1_xreal_0(sK1) <- (217, ~266) [trivial inequality removal 4583] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:12,thDist:-12}
4618. $false <- (185, 217, ~266) [subsumption resolution 4588,1822] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:12,thDist:-12}
4619. ~185 | ~217 | 266 [avatar contradiction clause 4618]
4626. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | v2_xxreal_0(sK2) | v3_xxreal_0(sK2) <- (3, 35, ~203) [subsumption resolution 4317,742] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4627. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | v2_xxreal_0(sK2) | v3_xxreal_0(sK2) <- (3, 35, 202, ~203) [subsumption resolution 4626,1986] {a:2,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4763. v1_xreal_0(k1_numbers) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (15, 41, 207) [superposition 2050,2605] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4790. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(sK2,k6_numbers) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 35, ~203) [subsumption resolution 4307,747] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4791. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(sK2,k6_numbers) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 5, 35, ~203) [subsumption resolution 4790,752] {a:2,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4792. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(sK2,k6_numbers) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (~2, 4, 5, 35, ~203) [subsumption resolution 4791,737] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4793. k6_numbers != k4_xxreal_3(sK2,k6_numbers) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (~2, 4, 5, 35, ~203, 252) [forward demodulation 4792,3499] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4795. 296 <=> k6_numbers = k4_xxreal_3(sK2,k6_numbers) [avatar definition]
4797. k6_numbers != k4_xxreal_3(sK2,k6_numbers) <- (~296) [avatar component clause 4795] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
4798. 167 | 166 | ~296 | 2 | ~4 | ~5 | ~35 | 203 | ~252 [avatar split clause 4793,3498,1989,905,750,745,735,4795,1703,1707]
4917. v7_ordinal1(sK1) <- (165) [resolution 1701,305] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
4924. 194 | ~165 [avatar split clause 4917,1699,1887]
4979. v1_xreal_0(sK1) <- (194) [resolution 1889,296] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
5047. ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (~205) [resolution 2041,493] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
5048. ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (~205) [resolution 2041,489] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
5192. ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) | ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (~210) [resolution 2081,489] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
5318. ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (208) [resolution 443,2055] {a:3,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
5325. ~v3_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (202, 208) [subsumption resolution 5318,1986] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
5326. ~206 | ~202 | ~208 [avatar split clause 5325,2053,1985,2044]
5701. ~v2_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (~240) [resolution 493,3073] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
5792. v1_xreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (37, 207) [superposition 2050,914] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
5896. v1_xreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (38, 244) [superposition 3093,918] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
6039. v1_xreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (39, 207) [superposition 2050,922] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6269. v1_xreal_0(k1_numbers) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (15, 43, 207) [superposition 2050,2607] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6555. k6_numbers != k6_numbers | ~v1_xreal_0(sK2) <- (217, ~296) [superposition 4797,2282] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6570. ~v1_xreal_0(sK2) <- (217, ~296) [trivial inequality removal 6555] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6572. $false <- (186, 217, ~296) [subsumption resolution 6570,1827] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6573. ~186 | ~217 | 296 [avatar contradiction clause 6572]
6782. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (15, 43, ~93, 207) [subsumption resolution 6269,1190] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6783. ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (4, 15, 43, ~93, 207) [subsumption resolution 6782,747] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6784. v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (4, 5, 15, 43, ~93, 207) [subsumption resolution 6783,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6785. ~v2_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (~2, 4, 5, 15, 43, ~93, 207) [subsumption resolution 6784,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6786. ~163 | ~166 | 2 | ~4 | ~5 | ~15 | ~43 | 93 | ~207 [avatar split clause 6785,2048,1188,937,800,750,745,735,1703,1690]
6824. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (~7, 39, 207) [subsumption resolution 6039,762] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6825. ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (4, ~7, 39, 207) [subsumption resolution 6824,747] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6826. v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (4, 5, ~7, 39, 207) [subsumption resolution 6825,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6827. ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (~2, 4, 5, ~7, 39, 207) [subsumption resolution 6826,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6828. ~164 | ~166 | 2 | ~4 | ~5 | 7 | ~39 | ~207 [avatar split clause 6827,2048,921,760,750,745,735,1703,1694]
6900. ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (5, ~205) [subsumption resolution 5047,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
6901. ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (4, 5, ~205) [subsumption resolution 6900,747] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
6902. ~166 | ~163 | ~4 | ~5 | 205 [avatar split clause 6901,2040,750,745,1690,1703]
6921. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (15, 41, ~93, 207) [subsumption resolution 4763,1190] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6922. ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (4, 15, 41, ~93, 207) [subsumption resolution 6921,747] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6923. v1_xreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (4, 5, 15, 41, ~93, 207) [subsumption resolution 6922,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6924. ~v3_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (~2, 4, 5, 15, 41, ~93, 207) [subsumption resolution 6923,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
6925. ~164 | ~167 | 2 | ~4 | ~5 | ~15 | ~41 | 93 | ~207 [avatar split clause 6924,2048,1188,929,800,750,745,735,1707,1694]
6939. ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (5, ~205) [subsumption resolution 5048,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
6940. ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (4, 5, ~205) [subsumption resolution 6939,747] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
6941. ~167 | ~164 | ~4 | ~5 | 205 [avatar split clause 6940,2040,750,745,1694,1707]
7049. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (~7, 37, 207) [subsumption resolution 5792,762] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
7050. ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (4, ~7, 37, 207) [subsumption resolution 7049,747] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
7051. v1_xreal_0(sK0) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (4, 5, ~7, 37, 207) [subsumption resolution 7050,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
7052. ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (~2, 4, 5, ~7, 37, 207) [subsumption resolution 7051,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
7053. ~163 | ~167 | 2 | ~4 | ~5 | 7 | ~37 | ~207 [avatar split clause 7052,2048,913,760,750,745,735,1707,1690]
7113. ~v3_xxreal_0(sK1) | ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (4, ~210) [subsumption resolution 5192,747] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
7114. ~v3_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (4, 5, ~210) [subsumption resolution 7113,752] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
7115. ~164 | ~167 | ~4 | ~5 | 210 [avatar split clause 7114,2080,750,745,1707,1694]
7192. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (~7, 38, 244) [subsumption resolution 5896,762] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
7193. ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (5, ~7, 38, 244) [subsumption resolution 7192,752] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
7194. v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (3, 5, ~7, 38, 244) [subsumption resolution 7193,742] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
7195. ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (~2, 3, 5, ~7, 38, 244) [subsumption resolution 7194,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
7196. ~169 | ~164 | 2 | ~3 | ~5 | 7 | ~38 | ~244 [avatar split clause 7195,3091,917,760,750,740,735,1694,1716]
7278. ~v2_xxreal_0(sK2) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (3, ~240) [subsumption resolution 5701,742] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
7279. ~v2_xxreal_0(sK2) | ~v2_xxreal_0(sK0) <- (3, 5, ~240) [subsumption resolution 7278,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
7280. ~163 | ~169 | ~3 | ~5 | 240 [avatar split clause 7279,3072,750,740,1716,1690]
7326. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (15, 44, ~93, 244) [subsumption resolution 3719,1190] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
7327. ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (5, 15, 44, ~93, 244) [subsumption resolution 7326,752] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
7328. v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (3, 5, 15, 44, ~93, 244) [subsumption resolution 7327,742] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
7329. ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK2) <- (~2, 3, 5, 15, 44, ~93, 244) [subsumption resolution 7328,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
7330. ~169 | ~163 | 2 | ~3 | ~5 | ~15 | ~44 | 93 | ~244 [avatar split clause 7329,3091,1188,941,800,750,740,735,1690,1716]
7622. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (3, 39, ~203) [subsumption resolution 4320,742] {a:2,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
7623. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (3, 39, 202, ~203) [subsumption resolution 7622,1986] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
7624. ~206 | ~169 | 207 | ~253 | ~3 | ~39 | ~202 | 203 [avatar split clause 7623,1989,1985,921,740,3543,2048,1716,2044]
7642. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (38, ~201, 214) [subsumption resolution 4289,2222] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
7643. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (3, 38, ~201, 214) [subsumption resolution 7642,742] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
7644. ~169 | ~213 | 212 | ~253 | ~3 | ~38 | 201 | ~214 [avatar split clause 7643,2221,1979,917,740,3543,2130,2162,1716]
7659. 240 | 241 | 244 | ~242 [avatar split clause 3686,3080,3091,3076,3072]
7666. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | v1_xreal_0(sK2) | v3_xxreal_0(k4_xxreal_3(sK1,sK0)) | v2_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (3, 24, ~203) [subsumption resolution 4316,742] {a:2,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
7667. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK2) | v3_xxreal_0(k4_xxreal_3(sK1,sK0)) | v2_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (3, 24, 202, ~203) [subsumption resolution 7666,1986] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
7668. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK2) | v3_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (3, 24, 202, ~203, ~205) [subsumption resolution 7667,2041] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
7669. 206 | 186 | ~252 | ~3 | ~24 | ~202 | 203 | 205 [avatar split clause 7668,2040,1989,1985,843,740,3498,1825,2044]
7774. k1_numbers != k4_xxreal_3(k4_xxreal_3(sK1,sK0),sK2) <- (~1, 235) [backward demodulation 732,2929] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
7813. 341 <=> k1_numbers = k4_xxreal_3(k4_xxreal_3(sK1,sK0),sK2) [avatar definition]
7815. k1_numbers != k4_xxreal_3(k4_xxreal_3(sK1,sK0),sK2) <- (~341) [avatar component clause 7813] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
7816. ~341 | 1 | ~235 [avatar split clause 7774,2928,730,7813]
7893. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) <- (~1, 38) [superposition 732,918] {a:1,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
7918. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) <- (~1, 38, 202) [subsumption resolution 7893,1986] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
7919. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(sK2) <- (~1, 3, 38, 202) [subsumption resolution 7918,742] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
7960. k6_numbers != k4_xxreal_3(sK1,k6_numbers) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (24, ~252) [superposition 3500,844] {a:2,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
7970. k6_numbers != k6_numbers | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | v3_xxreal_0(k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (24, ~252) [superposition 3500,844] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
7971. k6_numbers != k6_numbers | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (35, ~252) [superposition 3500,906] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
7980. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (35, ~252) [trivial inequality removal 7971] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
7981. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | v3_xxreal_0(k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (24, ~252) [trivial inequality removal 7970] {a:2,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8004. k2_xxreal_0 != k2_xxreal_0 | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (37, ~253) [superposition 3545,914] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8015. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (37, ~253) [trivial inequality removal 8004] {a:2,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8026. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 37, ~253) [subsumption resolution 8015,747] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8027. v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(sK1) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 37, 242, ~253) [subsumption resolution 8026,3081] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8028. v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 37, 167, 242, ~253) [subsumption resolution 8027,1709] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8029. ~240 | 244 | ~4 | ~37 | ~167 | ~242 | 253 [avatar split clause 8028,3543,3080,1707,913,745,3091,3072]
8041. k6_numbers != k4_xxreal_3(sK1,k6_numbers) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (5, 24, ~252) [subsumption resolution 7960,752] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8042. k6_numbers != k4_xxreal_3(sK1,k6_numbers) | v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (3, 5, 24, ~252) [subsumption resolution 8041,742] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8043. k6_numbers != k4_xxreal_3(sK1,k6_numbers) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (~2, 3, 5, 24, ~252) [subsumption resolution 8042,737] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8063. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | v3_xxreal_0(k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 24, ~252) [subsumption resolution 7981,747] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8064. v1_xreal_0(sK1) | v3_xxreal_0(k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 24, 242, ~252) [subsumption resolution 8063,3081] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
8065. 240 | 241 | 185 | ~4 | ~24 | ~242 | 252 [avatar split clause 8064,3498,3080,843,745,1820,3076,3072]
8110. 185 | ~194 [avatar split clause 4979,1887,1820]
8144. ~169 | ~206 | 207 | ~253 | 1 | ~3 | ~38 | ~202 [avatar split clause 7919,1985,917,740,730,3543,2048,2044,1716]
9388. ~v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~207) [superposition 2049,495] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
9390. ~v1_xreal_0(k6_numbers) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (35, ~207) [superposition 2049,906] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
9400. ~v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK1) <- (5, ~207) [subsumption resolution 9388,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
9401. ~v1_xreal_0(k4_xxreal_3(sK0,sK1)) <- (4, 5, ~207) [subsumption resolution 9400,747] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
9402. ~212 | ~4 | ~5 | 207 [avatar split clause 9401,2048,750,745,2130]
9461. k1_numbers != k1_numbers | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (15, 41, ~235) [superposition 2930,2605] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9471. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (15, 41, ~235) [trivial inequality removal 9461] {a:2,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9511. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 41, ~235) [subsumption resolution 9471,747] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9512. v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 41, ~235, 242) [subsumption resolution 9511,3081] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9513. v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 41, 167, ~235, 242) [subsumption resolution 9512,1709] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9514. ~241 | 244 | ~4 | ~15 | ~41 | ~167 | 235 | ~242 [avatar split clause 9513,3080,2928,1707,929,800,745,3091,3076]
9573. k1_numbers != k1_numbers | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (15, 42, ~235) [superposition 2930,2606] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9583. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (15, 42, ~235) [trivial inequality removal 9573] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9625. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(sK1) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 42, ~235) [subsumption resolution 9583,747] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9626. v1_xreal_0(sK1) | ~v3_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 42, ~235, 242) [subsumption resolution 9625,3081] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9627. ~v3_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 42, ~185, ~235, 242) [subsumption resolution 9626,1821] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9628. ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 15, 42, 167, ~185, ~235, 242) [subsumption resolution 9627,1709] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9629. ~241 | ~4 | ~15 | ~42 | ~167 | 185 | 235 | ~242 [avatar split clause 9628,3080,2928,1820,1707,933,800,745,3076]
9634. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (35, ~208) [superposition 2054,906] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
9651. v2_xxreal_0(k6_numbers) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (24, 240) [superposition 3074,844] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
9663. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (24, ~143, 240) [subsumption resolution 9651,1538] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
9664. ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (5, 24, ~143, 240) [subsumption resolution 9663,752] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
9665. v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (3, 5, 24, ~143, 240) [subsumption resolution 9664,742] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
9666. v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (~2, 3, 5, 24, ~143, 240) [subsumption resolution 9665,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
9667. v2_xxreal_0(sK2) <- (~2, 3, 5, 24, ~143, ~170, 240) [subsumption resolution 9666,1721] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
9668. $false <- (~2, 3, 5, 24, ~143, ~169, ~170, 240) [subsumption resolution 9667,1717] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
9669. 2 | ~3 | ~5 | ~24 | 143 | 169 | 170 | ~240 [avatar contradiction clause 9668]
9773. k6_numbers != k4_xxreal_3(sK1,k6_numbers) | v3_xxreal_0(sK2) <- (~2, 3, 5, 24, ~169, ~252) [subsumption resolution 8043,1717] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9774. 170 | ~266 | 2 | ~3 | ~5 | ~24 | 169 | 252 [avatar split clause 9773,3498,1716,843,750,740,735,4035,1720]
9793. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (15, 42, ~201, 214) [subsumption resolution 4293,2222] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
9794. k1_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (3, 15, 42, ~201, 214) [subsumption resolution 9793,742] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
9795. ~170 | ~213 | 212 | ~235 | ~3 | ~15 | ~42 | 201 | ~214 [avatar split clause 9794,2221,1979,933,800,740,2928,2130,2162,1720]
9796. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(sK2) | v3_xxreal_0(sK2) <- (3, 35, 202, ~203, ~207) [subsumption resolution 4627,2049] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
9797. k6_numbers != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v3_xxreal_0(sK2) <- (3, 35, ~169, 202, ~203, ~207) [subsumption resolution 9796,1717] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
9798. 170 | ~252 | ~3 | ~35 | 169 | ~202 | 203 | 207 [avatar split clause 9797,2048,1989,1985,1716,905,740,3498,1720]
9851. k1_numbers != k4_xxreal_3(k2_xxreal_0,sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (39, ~341) [superposition 7815,922] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
9890. 381 <=> k1_numbers = k4_xxreal_3(k2_xxreal_0,sK2) [avatar definition]
9891. k1_numbers = k4_xxreal_3(k2_xxreal_0,sK2) <- (381) [avatar component clause 9890] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
9892. k1_numbers != k4_xxreal_3(k2_xxreal_0,sK2) <- (~381) [avatar component clause 9890] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
9931. k1_numbers != k1_numbers | ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK2) | v1_xreal_0(k2_xxreal_0) | ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (15, 42, ~381) [superposition 9892,2606] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
9936. ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK2) | v1_xreal_0(k2_xxreal_0) | ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (15, 42, ~381) [trivial inequality removal 9931] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10013. v3_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (241) [resolution 3078,483] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
10014. v3_xxreal_0(k6_numbers) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (24, 241) [superposition 3078,844] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10030. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (24, ~142, 241) [subsumption resolution 10014,1533] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10031. ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (5, 24, ~142, 241) [subsumption resolution 10030,752] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10032. v1_xreal_0(sK0) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (3, 5, 24, ~142, 241) [subsumption resolution 10031,742] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10033. v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (~2, 3, 5, 24, ~142, 241) [subsumption resolution 10032,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10123. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (35, 90, ~208) [subsumption resolution 9634,1172] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10124. ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 35, 90, ~208) [subsumption resolution 10123,747] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10125. v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 5, 35, 90, ~208) [subsumption resolution 10124,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10126. v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (~2, 4, 5, 35, 90, ~208) [subsumption resolution 10125,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10127. 167 | 166 | 2 | ~4 | ~5 | ~35 | ~90 | 208 [avatar split clause 10126,2053,1170,905,750,745,735,1703,1707]
10128. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (35, 145, ~207) [subsumption resolution 9390,1549] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10129. ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 35, 145, ~207) [subsumption resolution 10128,747] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10130. v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 5, 35, 145, ~207) [subsumption resolution 10129,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10131. v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (~2, 4, 5, 35, 145, ~207) [subsumption resolution 10130,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10132. 167 | 166 | 2 | ~4 | ~5 | ~35 | ~145 | 207 [avatar split clause 10131,2048,1547,905,750,745,735,1703,1707]
10152. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 35, ~252) [subsumption resolution 7980,747] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10153. v1_xreal_0(k4_xxreal_3(sK0,sK2)) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 35, 242, ~252) [subsumption resolution 10152,3081] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10154. 167 | 166 | 244 | ~4 | ~35 | ~242 | 252 [avatar split clause 10153,3498,3080,905,745,3091,1703,1707]
10162. v3_xxreal_0(sK2) <- (~2, 3, 5, 24, ~142, ~169, 241) [subsumption resolution 10033,1717] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10163. 170 | 2 | ~3 | ~5 | ~24 | 142 | 169 | ~241 [avatar split clause 10162,3076,1716,1531,843,750,740,735,1720]
10164. v3_xxreal_0(sK2) | v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (3, 241) [subsumption resolution 10013,742] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
10165. v3_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (3, ~164, 241) [subsumption resolution 10164,1695] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
10166. v3_xxreal_0(sK2) <- (3, 5, ~164, 241) [subsumption resolution 10165,752] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
10167. 170 | ~3 | ~5 | 164 | ~241 [avatar split clause 10166,3076,1694,750,740,1720]
10175. ~v1_xxreal_0(sK2) | v1_xreal_0(k2_xxreal_0) | ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (12, 15, 42, ~381) [subsumption resolution 9936,787] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10176. v1_xreal_0(k2_xxreal_0) | ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (3, 12, 15, 42, ~381) [subsumption resolution 10175,742] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10177. ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (3, ~7, 12, 15, 42, ~381) [subsumption resolution 10176,762] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10178. ~v3_xxreal_0(sK2) <- (3, ~7, 9, 12, 15, 42, ~381) [subsumption resolution 10177,772] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10179. ~170 | ~3 | 7 | ~9 | ~12 | ~15 | ~42 | 381 [avatar split clause 10178,9890,933,800,785,770,760,740,1720]
10203. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (40, ~201, 214) [subsumption resolution 4291,2222] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
10204. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (3, 40, ~201, 214) [subsumption resolution 10203,742] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
10216. ~170 | ~210 | 212 | ~253 | ~3 | ~40 | 201 | ~214 [avatar split clause 10204,2221,1979,925,740,3543,2130,2080,1720]
10310. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~1) [superposition 732,495] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
10322. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(sK2,k4_xxreal_3(sK1,sK0)) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (~1) [superposition 732,495] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
10328. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v3_xxreal_0(sK2) <- (~1, 40) [superposition 732,926] {a:1,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
10336. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) | ~v1_xxreal_0(sK1) <- (~1, 5) [subsumption resolution 10310,752] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
10337. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) <- (~1, 4, 5) [subsumption resolution 10336,747] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
10338. ~201 | 1 | ~4 | ~5 [avatar split clause 10337,750,745,730,1979]
10364. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(sK2,k4_xxreal_3(sK1,sK0)) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK0)) <- (~1, 3) [subsumption resolution 10322,742] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
10365. k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) != k4_xxreal_3(sK2,k4_xxreal_3(sK1,sK0)) <- (~1, 3, 202) [subsumption resolution 10364,1986] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
10366. ~203 | 1 | ~3 | ~202 [avatar split clause 10365,1985,740,730,1989]
10372. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v3_xxreal_0(sK2) <- (~1, 40, 202) [subsumption resolution 10328,1986] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
10373. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK1,sK0)) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v3_xxreal_0(sK2) <- (~1, 3, 40, 202) [subsumption resolution 10372,742] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
10374. k2_xxreal_0 != k4_xxreal_3(sK1,k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK0)) | ~v3_xxreal_0(sK2) <- (~1, 3, 40, 202, ~207) [subsumption resolution 10373,2049] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
10375. ~170 | ~205 | ~253 | 1 | ~3 | ~40 | ~202 | 207 [avatar split clause 10374,2048,1985,925,740,730,3543,2040,1720]
10470. v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) | v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (205) [resolution 2042,485] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
10475. v2_xxreal_0(k6_numbers) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (35, 205) [superposition 2042,906] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10507. k2_xxreal_0 != k2_xxreal_0 | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (39, ~253) [superposition 3545,922] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10514. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (39, ~253) [trivial inequality removal 10507] {a:2,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10523. ~v1_xxreal_0(k4_xxreal_3(sK0,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 39, ~253) [subsumption resolution 10514,747] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10524. v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 39, 242, ~253) [subsumption resolution 10523,3081] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10525. v1_xreal_0(k4_xxreal_3(sK0,sK2)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 39, 166, 242, ~253) [subsumption resolution 10524,1705] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10526. v1_xreal_0(k4_xxreal_3(sK0,sK2)) <- (4, 39, 166, 241, 242, ~253) [subsumption resolution 10525,3078] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
10557. 244 | ~4 | ~39 | ~166 | ~241 | ~242 | 253 [avatar split clause 10526,3543,3080,3076,1703,921,745,3091]
10565. ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (~206) [resolution 2045,491] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
10627. v1_xreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (40, 244) [superposition 3093,926] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10629. v1_xreal_0(k1_numbers) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (15, 42, 244) [superposition 3093,2606] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10633. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (~7, 40, 244) [subsumption resolution 10627,762] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10634. ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (5, ~7, 40, 244) [subsumption resolution 10633,752] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10635. v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (3, 5, ~7, 40, 244) [subsumption resolution 10634,742] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10636. ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (~2, 3, 5, ~7, 40, 244) [subsumption resolution 10635,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10640. ~v2_xxreal_0(sK0) <- (~2, 3, 5, ~7, 40, 170, 244) [subsumption resolution 10636,1722] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10641. ~163 | 2 | ~3 | ~5 | 7 | ~40 | ~170 | ~244 [avatar split clause 10640,3091,1720,925,760,750,740,735,1690]
10660. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (15, 42, ~93, 244) [subsumption resolution 10629,1190] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10661. ~v1_xxreal_0(sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (5, 15, 42, ~93, 244) [subsumption resolution 10660,752] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10662. v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (3, 5, 15, 42, ~93, 244) [subsumption resolution 10661,742] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10663. ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK2) <- (~2, 3, 5, 15, 42, ~93, 244) [subsumption resolution 10662,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10664. ~v3_xxreal_0(sK0) <- (~2, 3, 5, 15, 42, ~93, 170, 244) [subsumption resolution 10663,1722] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
10665. ~164 | 2 | ~3 | ~5 | ~15 | ~42 | 93 | ~170 | ~244 [avatar split clause 10664,3091,1720,1188,933,800,750,740,735,1694]
10675. ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (5, ~206) [subsumption resolution 10565,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
10676. ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (4, 5, ~206) [subsumption resolution 10675,747] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
10677. ~166 | ~164 | ~4 | ~5 | 206 [avatar split clause 10676,2044,750,745,1694,1703]
10714. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (39, ~341, 381) [subsumption resolution 9851,9891] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
10715. ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (4, 39, ~341, 381) [subsumption resolution 10714,747] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
10716. v1_xreal_0(sK0) | ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (4, 5, 39, ~341, 381) [subsumption resolution 10715,752] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
10717. ~v2_xxreal_0(sK1) | ~v3_xxreal_0(sK0) <- (~2, 4, 5, 39, ~341, 381) [subsumption resolution 10716,737] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
10718. ~164 | ~166 | 2 | ~4 | ~5 | ~39 | 341 | ~381 [avatar split clause 10717,9890,7813,921,750,745,735,1703,1694]
10746. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (35, ~143, 205) [subsumption resolution 10475,1538] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10747. ~v1_xxreal_0(sK0) | v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 35, ~143, 205) [subsumption resolution 10746,747] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10748. v1_xreal_0(sK0) | v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (4, 5, 35, ~143, 205) [subsumption resolution 10747,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10749. v2_xxreal_0(sK1) | v3_xxreal_0(sK1) <- (~2, 4, 5, 35, ~143, 205) [subsumption resolution 10748,737] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
10750. 167 | 166 | 2 | ~4 | ~5 | ~35 | 143 | ~205 [avatar split clause 10749,2040,1536,905,750,745,735,1703,1707]
10751. v3_xxreal_0(sK0) | v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (5, 205) [subsumption resolution 10470,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
10752. v3_xxreal_0(sK0) | v2_xxreal_0(sK1) <- (4, 5, 205) [subsumption resolution 10751,747] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
10753. 166 | 164 | ~4 | ~5 | ~205 [avatar split clause 10752,2040,750,745,1694,1703]
10777. $false [avatar sat refutation 733,738,743,748,753,758,763,768,773,783,788,798,803,808,845,907,915,919,923,927,931,935,939,943,1173,1191,1197,1409,1529,1534,1539,1550,1779,1781,1818,1823,2027,2056,2113,2201,2289,3020,3104,3107,3178,3263,3423,3439,3473,3504,3530,4357,4619,4798,4924,5326,6573,6786,6828,6902,6925,6941,7053,7115,7196,7280,7330,7624,7644,7659,7669,7816,8029,8065,8110,8144,9402,9514,9629,9669,9774,9795,9798,10127,10132,10154,10163,10167,10179,10216,10338,10366,10375,10557,10641,10665,10677,10718,10750,10753] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:17548,thDist:-17548}
% SZS output end Proof for xxreal_3__l92_xxreal_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 108
% >>> Preprocessing
% Introduced skolems: 26
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 270
% Generated clauses: 9204
% Active clauses: 476
% Passive clauses: 987
% Final active clauses: 283
% Final passive clauses: 342
% Main loop iterations started: 569
% 
% >>> Simplifying Inferences
% Duplicate literals: 1677
% Trivial inequalities: 115
% Fw subsumption resolutions: 2577
% Fw demodulations: 64
% Bw demodulations: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 403
% Forward subsumptions: 3002
% Backward subsumptions: 51
% 
% >>> Generating Inferences
% Binary resolution: 570
% Unit resulting resolution: 38
% Forward superposition: 1688
% Backward superposition: 1439
% Self superposition: 4
% Equality factoring: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 722
% Split components: 1934
% Sat splitting refutations: 87
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 1329
% SAT solver unit clauses: 68
% SAT solver binary clauses: 217
% 
% Memory used [KB]: 11129
% Time elapsed: 0.063 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% SAT solver time: 0.003 s
% minimizing solver time: 0.002 s
% sat proof minimization: 0.001 s
% term sharing: 0.005 s
% congruence closure: 0.009 s ( own 0.007 s ) 
% binary resolution index maintenance: 0.001 s
% backward superposition index maintenance: 0.001 s
% forward subsumption: 0.012 s
% forward subsumption resolution: 0.003 s
% backward subsumption: 0.001 s
% forward demodulation: 0.002 s
% unit resulting resolution: 0.010 s ( own 0.007 s ) 
% superposition: 0.002 s ( own 0.001 s ) 
% passive container maintenance: 0.001 s
% other: 0.063 s ( own 0.015 s ) 

