i: [1,236,0,2,1,1,0,0] t5_vectsp_2
a: [3,364,0,2,1,236]
s: 364
i: [1,235,0,2,1,1,0,0] t5_vectsp_2
a: [3,369,0,2,1,235]
s: 369
i: [1,234,0,2,1,1,0,0] t5_vectsp_2
a: [3,374,0,2,1,234]
s: 374
i: [1,233,0,2,1,1,0,0] t5_vectsp_2
a: [3,379,0,2,1,233]
s: 379
i: [1,232,0,2,1,1,0,0] t5_vectsp_2
a: [3,384,0,2,1,232]
s: 384
i: [1,231,0,2,1,1,0,0] t5_vectsp_2
a: [3,389,0,2,1,231]
s: 389
i: [1,230,0,2,1,1,0,0] t5_vectsp_2
a: [3,394,0,2,1,230]
s: 394
i: [1,240,0,8,2,1,0,0] t5_vectsp_2
a: [3,344,0,4,1,240]
s: 344
i: [1,238,0,4,1,1,0,0] t5_vectsp_2
a: [3,354,0,4,1,238]
s: 354
i: [1,237,0,4,1,1,0,0] t5_vectsp_2
a: [3,359,0,4,1,237]
s: 359
i: [1,239,0,7,1,1,0,0] t5_vectsp_2
a: [3,349,0,7,1,239]
s: 349
i: [1,241,0,2,1,0,0,255] spc1_boole
a: [3,399,0,2,1,241]
s: 399
i: [1,242,0,2,1,0,0,255] fc1_xboole_0
a: [3,404,0,2,1,242]
s: 404
i: [1,322,0,2,1,0,0,2] rc2_xboole_0
a: [3,409,0,2,1,322]
s: 409
i: [1,323,0,2,1,0,0,255] existence_l1_algstr_0
a: [3,414,0,2,1,323]
s: 414
i: [1,324,0,2,1,0,0,255] existence_l4_algstr_0
a: [3,419,0,2,1,324]
s: 419
i: [1,325,0,2,1,0,0,255] existence_l5_algstr_0
a: [3,424,0,2,1,325]
s: 424
i: [1,326,0,2,1,0,0,255] existence_l2_algstr_0
a: [3,429,0,2,1,326]
s: 429
i: [1,327,0,2,1,0,0,2] existence_l4_struct_0
a: [3,434,0,2,1,327]
s: 434
i: [1,328,0,2,1,0,0,2] existence_l3_algstr_0
a: [3,439,0,2,1,328]
s: 439
i: [1,329,0,2,1,0,0,1] existence_l6_algstr_0
a: [3,444,0,2,1,329]
s: 444
i: [1,330,0,2,1,0,0,2] existence_l3_struct_0
a: [3,449,0,2,1,330]
s: 449
i: [1,331,0,2,1,0,0,2] existence_l2_struct_0
a: [3,454,0,2,1,331]
s: 454
i: [1,332,0,2,1,0,0,2] rc1_xboole_0
a: [3,459,0,2,1,332]
s: 459
i: [1,333,0,2,1,0,0,3] existence_l1_struct_0
a: [3,464,0,2,1,333]
s: 464
i: [1,336,0,2,1,0,0,2] rc20_struct_0
a: [3,469,0,2,1,336]
s: 469
i: [1,335,0,2,1,0,0,2] rc20_struct_0
a: [3,474,0,2,1,335]
s: 474
i: [1,334,0,2,1,0,0,2] rc20_struct_0
a: [3,479,0,2,1,334]
s: 479
i: [1,317,0,3,1,0,0,255] reflexivity_r1_tarski
s: 317
i: [1,243,0,4,2,0,0,255] dt_l1_algstr_0
s: 243
i: [1,244,0,4,2,0,0,255] dt_l5_algstr_0
s: 244
d: [2,481,1,2,1,1,40,243,414]
a: [3,485,1,2,1,481]
s: 485
i: [1,245,0,4,2,0,0,255] dt_l5_algstr_0
s: 245
d: [2,487,1,2,1,1,40,244,424]
a: [3,491,1,2,1,487]
s: 491
i: [1,246,0,4,2,0,0,255] dt_l4_algstr_0
s: 246
i: [1,247,0,4,2,0,0,255] dt_l4_algstr_0
s: 247
d: [2,493,1,2,1,1,40,245,424]
a: [3,497,1,2,1,493]
s: 497
i: [1,248,0,4,2,0,0,255] dt_l2_algstr_0
s: 248
d: [2,499,1,2,1,1,40,246,419]
a: [3,504,1,2,1,499]
s: 504
i: [1,249,0,4,2,0,0,255] dt_l2_algstr_0
s: 249
i: [1,250,0,4,2,0,0,2] dt_l4_struct_0
s: 250
d: [2,511,1,2,1,1,40,247,419]
a: [3,516,1,2,1,511]
s: 516
i: [1,251,0,4,2,0,0,2] dt_l4_struct_0
s: 251
d: [2,523,1,2,1,1,40,248,429]
a: [3,527,1,2,1,523]
s: 527
i: [1,253,0,4,2,0,0,2] dt_l3_algstr_0
s: 253
i: [1,258,0,4,2,0,0,255] dt_l6_algstr_0
s: 258
d: [2,567,1,2,1,1,40,258,364]
a: [3,572,1,2,1,567]
s: 572
i: [1,259,0,4,2,0,0,255] dt_l6_algstr_0
s: 259
d: [2,591,1,2,1,1,40,259,364]
a: [3,596,1,2,1,591]
s: 596
i: [1,265,0,4,2,0,0,2] dt_l3_struct_0
s: 265
i: [1,269,0,4,2,0,0,2] dt_l2_struct_0
s: 269
d: [2,579,2,2,1,1,40,572,249]
a: [3,584,2,2,1,579]
s: 584
i: [1,316,0,4,1,0,0,1] existence_m1_subset_1
s: 316
d: [2,580,2,2,1,1,40,572,248]
a: [3,589,2,2,1,580]
s: 589
i: [1,254,0,5,2,0,0,255] dt_u2_algstr_0
s: 254
i: [1,276,0,5,2,0,0,255] t6_boole
s: 276
d: [2,603,2,2,1,1,40,596,245]
a: [3,608,2,2,1,603]
s: 608
i: [1,313,0,5,2,0,0,255] t7_boole
s: 313
d: [2,604,2,2,1,1,40,596,244]
a: [3,613,2,2,1,604]
s: 613
i: [1,252,0,6,3,0,0,1] cc3_struct_0
s: 252
i: [1,270,0,6,2,0,0,2] fc10_struct_0
s: 270
r: 344
d: [2,641,3,2,1,1,40,584,243]
a: [3,645,3,2,1,641]
s: 645
a: [3,340,0,4,1,240]
s: 340
r: 340
d: [2,651,3,2,1,1,40,608,251]
a: [3,656,3,2,1,651]
s: 656
a: [3,339,0,4,1,240]
s: 339
s: 344
d: [2,658,3,2,1,1,40,613,247]
a: [3,663,3,2,1,658]
s: 663
i: [1,275,0,6,2,0,0,2] rc17_struct_0
s: 275
d: [2,668,3,4,2,1,40,252,608]
d: [2,687,3,2,1,2,21,668,389]
a: [3,691,3,2,1,687]
s: 691
i: [1,277,0,6,3,0,0,255] cc7_struct_0
s: 277
i: [1,278,0,6,3,0,0,255] cc5_struct_0
s: 278
d: [2,693,1,5,2,1,44,270,344]
d: [2,694,1,3,1,2,21,693,589]
a: [3,698,1,3,1,694]
s: 698
i: [1,280,0,6,3,0,0,2] cc17_struct_0
s: 280
d: [2,708,1,3,1,2,44,339,344]
a: [3,713,1,3,1,708]
s: 713
i: [1,281,0,6,3,0,0,2] cc2_struct_0
s: 281
i: [1,282,0,6,3,0,0,255] cc6_struct_0
s: 282
d: [2,529,1,2,1,1,40,249,429]
a: [3,533,1,2,1,529]
s: 533
i: [1,283,0,6,3,0,0,2] cc1_struct_0
s: 283
d: [2,535,1,2,1,1,40,250,434]
a: [3,540,1,2,1,535]
s: 540
i: [1,286,0,6,3,0,0,255] cc4_struct_0
s: 286
i: [1,309,0,6,2,0,0,255] t1_subset
s: 309
d: [2,547,1,2,1,1,40,251,434]
a: [3,552,1,2,1,547]
s: 552
i: [1,318,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 318
d: [2,555,1,2,1,1,40,253,439]
a: [3,560,1,2,1,555]
s: 560
i: [1,266,0,7,2,0,0,255] d7_struct_0
s: 266
i: [1,267,0,7,2,0,0,255] dt_u3_struct_0
s: 267
d: [2,568,1,2,1,1,40,258,444]
a: [3,577,1,2,1,568]
s: 577
i: [1,268,0,7,2,0,0,1] dt_k5_struct_0
s: 268
d: [2,592,1,2,1,1,40,259,444]
a: [3,601,1,2,1,592]
s: 601
i: [1,271,0,7,2,0,0,255] d6_struct_0
s: 271
i: [1,272,0,7,2,0,0,255] dt_u2_struct_0
s: 272
d: [2,615,1,2,1,1,40,265,449]
a: [3,620,1,2,1,615]
s: 620
i: [1,273,0,7,2,0,0,1] dt_k4_struct_0
s: 273
d: [2,623,1,2,1,1,40,269,454]
a: [3,629,1,2,1,623]
s: 629
i: [1,274,0,7,2,0,0,2] rc17_struct_0
s: 274
i: [1,284,0,7,3,0,0,255] cc8_struct_0
s: 284
d: [2,624,1,2,1,1,40,269,479]
a: [3,634,1,2,1,624]
s: 634
i: [1,287,0,7,3,0,0,255] cc11_struct_0
s: 287
d: [2,666,1,4,2,1,40,252,434]
a: [3,672,1,2,1,666]
s: 672
i: [1,288,0,7,3,0,0,255] cc11_struct_0
s: 288
i: [1,289,0,7,3,0,0,255] cc9_struct_0
s: 289
d: [2,718,1,4,2,1,40,277,469]
a: [3,722,1,2,1,718]
s: 722
i: [1,292,0,7,3,0,0,255] fc9_struct_0
s: 292
d: [2,724,1,4,2,1,40,278,464]
a: [3,730,1,2,1,724]
s: 730
i: [1,295,0,7,3,0,0,255] fc6_struct_0
s: 295
i: [1,297,0,7,3,0,0,255] rc22_struct_0
s: 297
d: [2,500,2,2,1,1,40,246,491]
a: [3,509,2,2,1,500]
s: 509
i: [1,300,0,7,3,0,0,2] fc2_struct_0
s: 300
d: [2,512,2,2,1,1,40,247,491]
a: [3,521,2,2,1,512]
s: 521
i: [1,302,0,7,3,0,0,255] rc4_struct_0
s: 302
i: [1,304,0,7,3,0,0,255] rc21_struct_0
s: 304
d: [2,536,2,2,1,1,40,250,497]
a: [3,545,2,2,1,536]
s: 545
i: [1,305,0,7,3,0,0,255] rc21_struct_0
s: 305
d: [2,556,2,2,1,1,40,253,516]
a: [3,565,2,2,1,556]
s: 565
i: [1,306,0,7,3,0,0,255] fc8_struct_0
s: 306
i: [1,307,0,7,3,0,0,255] fc7_struct_0
s: 307
d: [2,625,2,2,1,1,40,269,527]
a: [3,639,2,2,1,625]
s: 639
i: [1,308,0,7,3,0,0,2] fc1_struct_0
s: 308
d: [2,667,2,4,2,1,40,252,497]
a: [3,681,2,2,1,667]
s: 681
i: [1,311,0,7,2,0,0,255] t3_subset
s: 311
i: [1,312,0,7,3,0,0,1] t8_boole
s: 312
d: [2,725,2,4,2,1,40,278,485]
a: [3,739,2,2,1,725]
s: 739
i: [1,263,0,8,4,0,0,1] cc1_vectsp_1
s: 263
d: [2,899,1,6,3,1,40,263,364]
d: [2,901,1,4,2,2,21,899,379]
d: [2,902,1,2,1,3,21,901,394]
a: [3,906,1,2,1,902]
s: 906
i: [1,264,0,8,4,0,0,1] cc1_vectsp_1
s: 264
i: [1,291,0,8,3,0,0,1] fc11_struct_0
s: 291
d: [2,921,1,6,3,1,40,264,364]
d: [2,923,1,4,2,2,21,921,379]
d: [2,924,1,2,1,3,21,923,394]
a: [3,928,1,2,1,924]
s: 928
i: [1,294,0,8,3,0,0,2] rc19_struct_0
s: 294
d: [2,900,1,6,3,1,40,263,444]
a: [3,919,1,2,1,900]
s: 919
i: [1,310,0,8,3,0,0,255] t2_subset
s: 310
i: [1,290,0,9,4,0,0,255] cc10_struct_0
s: 290
d: [2,935,1,7,2,1,40,310,354]
a: [3,951,1,3,1,935]
s: 951
d: [2,962,4,0,0,3,52,394,645,951,300]
e: 962
d: [2,963,2,4,2,1,40,951,300]
d: [2,965,2,2,1,2,21,963,645]
d: [2,966,2,0,0,3,21,965,394]
e: 966
r: 951
i: [1,293,0,9,3,0,0,2] rc19_struct_0
s: 293
a: [3,950,1,3,1,935]
s: 950
i: [1,315,0,9,3,0,0,255] t5_subset
s: 315
i: [1,261,0,10,5,0,0,1] cc2_vectsp_1
s: 261
a: [3,947,1,4,1,935]
s: 947
i: [1,296,0,10,3,0,0,255] rc22_struct_0
s: 296
d: [2,936,1,7,2,1,40,310,359]
a: [3,956,1,4,1,936]
s: 956
i: [1,301,0,10,3,0,0,255] rc4_struct_0
s: 301
i: [1,303,0,10,3,0,0,255] rc21_struct_0
s: 303
d: [2,973,2,4,1,1,40,947,318]
a: [3,979,2,4,1,973]
s: 979
i: [1,314,0,10,3,0,0,255] t4_subset
s: 314
d: [2,983,2,4,1,1,40,956,318]
a: [3,989,2,4,1,983]
s: 989
i: [1,255,0,12,2,0,0,255] dt_u2_algstr_0
s: 255
i: [1,256,0,14,2,0,0,255] dt_u2_algstr_0
s: 256
d: [2,760,2,2,1,1,40,540,269]
a: [3,764,2,2,1,760]
s: 764
i: [1,319,0,17,4,0,0,1] dt_k6_algstr_0
s: 319
d: [2,1017,1,15,4,1,44,319,349]
d: [2,1018,1,11,3,2,21,1017,354]
d: [2,1019,1,7,2,3,21,1018,359]
d: [2,1020,1,5,1,4,21,1019,663]
a: [3,1024,1,5,1,1020]
s: 1024
i: [1,257,0,22,4,0,0,255] d18_algstr_0
s: 257
i: [1,298,0,24,8,0,0,1] t6_vectsp_1
s: 298
d: [2,1026,2,8,2,1,40,1024,310]
d: [2,1027,2,5,1,2,21,1026,950]
a: [3,1031,2,5,1,1027]
s: 1031
i: [1,299,0,24,8,0,0,1] t7_vectsp_1
s: 299
d: [2,1042,3,5,1,1,40,1031,318]
a: [3,1048,3,5,1,1042]
s: 1048
i: [1,320,0,29,6,0,0,255] dt_k5_binop_1
s: 320
i: [1,321,0,32,6,0,0,255] redefinition_k5_binop_1
s: 321
d: [2,768,2,4,2,1,40,560,280]
a: [3,773,2,2,1,768]
s: 773
d: [2,1033,1,22,8,1,44,298,344]
d: [2,1036,1,20,7,2,21,1033,364]
d: [2,1037,1,18,6,3,21,1036,906]
d: [2,1038,1,16,5,4,21,1037,369]
d: [2,1039,1,14,4,5,21,1038,374]
d: [2,1040,1,12,3,6,21,1039,384]
d: [2,1041,1,10,2,7,21,1040,394]
s: 1041
d: [2,1074,2,8,2,8,44,1041,349]
d: [2,1080,2,4,1,9,21,1074,359]
a: [3,1084,2,4,1,1080]
s: 1084
r: 1048
r: 1031
r: 1024
r: 349
d: [2,1092,3,7,3,1,44,291,1084]
d: [2,1099,3,4,2,2,21,1092,698]
d: [2,1100,3,2,1,3,21,1099,608]
d: [2,1101,3,0,0,4,21,1100,389]
e: 1101
r: 339
r: 1041
r: 344
r: 1084
s: 340
a: [3,343,0,4,1,240]
s: 343
d: [2,700,1,5,2,1,44,270,340]
d: [2,701,1,3,1,2,21,700,589]
a: [3,705,1,3,1,701]
s: 705
s: 349
a: [3,1083,2,4,1,1080]
s: 1083
s: 1024
d: [2,1103,1,22,8,1,44,299,340]
d: [2,1107,1,20,7,2,21,1103,364]
d: [2,1108,1,18,6,3,21,1107,928]
d: [2,1109,1,16,5,4,21,1108,369]
d: [2,1110,1,14,4,5,21,1109,374]
d: [2,1111,1,12,3,6,21,1110,384]
d: [2,1112,1,10,2,7,21,1111,394]
s: 1112
r: 698
r: 713
s: 713
d: [2,1091,0,6,1,2,23,349,1084]
a: [3,1096,0,6,1,1091]
s: 1096
d: [2,1130,2,8,2,8,44,1112,349]
d: [2,1136,2,4,1,9,21,1130,354]
a: [3,1140,2,4,1,1136]
s: 1140
r: 1024
r: 1083
r: 349
d: [2,1152,3,7,3,1,44,291,1140]
d: [2,1159,3,4,2,2,21,1152,705]
d: [2,1160,3,2,1,3,21,1159,608]
d: [2,1161,3,0,0,4,21,1160,389]
e: 1161
d: [2,709,1,5,2,1,44,270,344]
d: [2,715,1,3,1,2,21,709,589]
d: [2,1076,2,8,2,8,44,349,1041]
d: [2,1086,2,4,1,9,21,1076,359]
d: [2,1106,1,5,2,1,44,270,340]
d: [2,1119,1,3,1,2,21,1106,589]
d: [2,1132,2,8,2,8,44,349,1112]
d: [2,1142,2,4,1,9,21,1132,354]
f: 1161,1142,1119,1101,1086,924,902,715,603,591,580,567,230,231,232,233,234,235,236,237,238,239,240
% Refutation found. Thanks to Tanya!
% SZS status Theorem for vectsp_2__t5_vectsp_2
% SZS output start Proof for vectsp_2__t5_vectsp_2
1. ! [X0] : ((l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v5_vectsp_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k6_algstr_0(X0,X1,X2) = k5_struct_0(X0) => (k4_struct_0(X0) != X2 & k4_struct_0(X0) != X1))))) [input t5_vectsp_2]
2. ~! [X0] : ((l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v5_vectsp_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k6_algstr_0(X0,X1,X2) = k5_struct_0(X0) => (k4_struct_0(X0) != X2 & k4_struct_0(X0) != X1))))) [negated conjecture 1]
30. ! [X0] : ((l4_struct_0(X0) & ~v6_struct_0(X0)) => ~v9_struct_0(k5_struct_0(X0),X0)) [input fc11_struct_0]
56. ! [X0] : (l2_algstr_0(X0) => (l1_algstr_0(X0) & l2_struct_0(X0))) [input dt_l2_algstr_0]
60. ! [X0] : (l5_algstr_0(X0) => (l4_struct_0(X0) & l4_algstr_0(X0))) [input dt_l5_algstr_0]
68. ! [X0] : (l2_struct_0(X0) => v9_struct_0(k4_struct_0(X0),X0)) [input fc10_struct_0]
76. ! [X0] : (l6_algstr_0(X0) => (l5_algstr_0(X0) & l2_algstr_0(X0))) [input dt_l6_algstr_0]
79. ! [X0] : (l6_algstr_0(X0) => ((v5_vectsp_1(X0) & ~v2_struct_0(X0)) => (v2_vectsp_1(X0) & v1_vectsp_1(X0) & ~v2_struct_0(X0)))) [input cc1_vectsp_1]
84. ! [X0] : ((l6_algstr_0(X0) & v1_vectsp_1(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v13_algstr_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k4_struct_0(X0) = k6_algstr_0(X0,X1,k4_struct_0(X0)))) [input t6_vectsp_1]
85. ! [X0] : ((l6_algstr_0(X0) & v2_vectsp_1(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v13_algstr_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k4_struct_0(X0) = k6_algstr_0(X0,k4_struct_0(X0),X1))) [input t7_vectsp_1]
96. ? [X0] : (? [X1] : (? [X2] : (((k4_struct_0(X0) = X2 | k4_struct_0(X0) = X1) & k6_algstr_0(X0,X1,X2) = k5_struct_0(X0)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v5_vectsp_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
97. ? [X0] : (? [X1] : (? [X2] : ((k4_struct_0(X0) = X2 | k4_struct_0(X0) = X1) & k6_algstr_0(X0,X1,X2) = k5_struct_0(X0) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v5_vectsp_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0)) [flattening 96]
99. ! [X0] : ((l4_struct_0(X0) & l4_algstr_0(X0)) | ~l5_algstr_0(X0)) [ennf transformation 60]
101. ! [X0] : ((l1_algstr_0(X0) & l2_struct_0(X0)) | ~l2_algstr_0(X0)) [ennf transformation 56]
108. ! [X0] : ((l5_algstr_0(X0) & l2_algstr_0(X0)) | ~l6_algstr_0(X0)) [ennf transformation 76]
111. ! [X0] : (((v2_vectsp_1(X0) & v1_vectsp_1(X0) & ~v2_struct_0(X0)) | (~v5_vectsp_1(X0) | v2_struct_0(X0))) | ~l6_algstr_0(X0)) [ennf transformation 79]
112. ! [X0] : ((v2_vectsp_1(X0) & v1_vectsp_1(X0) & ~v2_struct_0(X0)) | ~v5_vectsp_1(X0) | v2_struct_0(X0) | ~l6_algstr_0(X0)) [flattening 111]
118. ! [X0] : (v9_struct_0(k4_struct_0(X0),X0) | ~l2_struct_0(X0)) [ennf transformation 68]
146. ! [X0] : (~v9_struct_0(k5_struct_0(X0),X0) | (~l4_struct_0(X0) | v6_struct_0(X0))) [ennf transformation 30]
147. ! [X0] : (~v9_struct_0(k5_struct_0(X0),X0) | ~l4_struct_0(X0) | v6_struct_0(X0)) [flattening 146]
156. ! [X0] : (! [X1] : (k4_struct_0(X0) = k6_algstr_0(X0,X1,k4_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l6_algstr_0(X0) | ~v1_vectsp_1(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 84]
157. ! [X0] : (! [X1] : (k4_struct_0(X0) = k6_algstr_0(X0,X1,k4_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l6_algstr_0(X0) | ~v1_vectsp_1(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0)) [flattening 156]
158. ! [X0] : (! [X1] : (k4_struct_0(X0) = k6_algstr_0(X0,k4_struct_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l6_algstr_0(X0) | ~v2_vectsp_1(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 85]
159. ! [X0] : (! [X1] : (k4_struct_0(X0) = k6_algstr_0(X0,k4_struct_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l6_algstr_0(X0) | ~v2_vectsp_1(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0)) [flattening 158]
188. ? [X0] : (? [X1] : (? [X2] : ((k4_struct_0(X0) = X2 | k4_struct_0(X0) = X1) & k6_algstr_0(X0,X1,X2) = k5_struct_0(X0) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v5_vectsp_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : ((k4_struct_0(sK0) = X2 | k4_struct_0(sK0) = X1) & k6_algstr_0(sK0,X1,X2) = k5_struct_0(sK0) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l6_algstr_0(sK0) & v4_rlvect_1(sK0) & v3_rlvect_1(sK0) & v5_vectsp_1(sK0) & v13_algstr_0(sK0) & ~v6_struct_0(sK0) & ~v2_struct_0(sK0)) [choice axiom]
189. ? [X1] : (? [X2] : ((k4_struct_0(sK0) = X2 | k4_struct_0(sK0) = X1) & k6_algstr_0(sK0,X1,X2) = k5_struct_0(sK0) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : ((k4_struct_0(sK0) = X2 | k4_struct_0(sK0) = sK1) & k5_struct_0(sK0) = k6_algstr_0(sK0,sK1,X2) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
190. ? [X2] : ((k4_struct_0(sK0) = X2 | k4_struct_0(sK0) = sK1) & k5_struct_0(sK0) = k6_algstr_0(sK0,sK1,X2) & m1_subset_1(X2,u1_struct_0(sK0))) => ((k4_struct_0(sK0) = sK2 | k4_struct_0(sK0) = sK1) & k5_struct_0(sK0) = k6_algstr_0(sK0,sK1,sK2) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
191. (((k4_struct_0(sK0) = sK2 | k4_struct_0(sK0) = sK1) & k5_struct_0(sK0) = k6_algstr_0(sK0,sK1,sK2) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l6_algstr_0(sK0) & v4_rlvect_1(sK0) & v3_rlvect_1(sK0) & v5_vectsp_1(sK0) & v13_algstr_0(sK0) & ~v6_struct_0(sK0) & ~v2_struct_0(sK0) [skolemisation 97,190,189,188]
230. ~v2_struct_0(sK0) [cnf transformation 191] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
231. ~v6_struct_0(sK0) [cnf transformation 191] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
232. v13_algstr_0(sK0) [cnf transformation 191] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
233. v5_vectsp_1(sK0) [cnf transformation 191] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
234. v3_rlvect_1(sK0) [cnf transformation 191] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
235. v4_rlvect_1(sK0) [cnf transformation 191] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
236. l6_algstr_0(sK0) [cnf transformation 191] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
237. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 191] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
238. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 191] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
239. k5_struct_0(sK0) = k6_algstr_0(sK0,sK1,sK2) [cnf transformation 191] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
240. k4_struct_0(sK0) = sK2 | k4_struct_0(sK0) = sK1 [cnf transformation 191] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
245. ~l5_algstr_0(X0) | l4_struct_0(X0) [cnf transformation 99] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
248. ~l2_algstr_0(X0) | l2_struct_0(X0) [cnf transformation 101] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
258. ~l6_algstr_0(X0) | l2_algstr_0(X0) [cnf transformation 108] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
259. ~l6_algstr_0(X0) | l5_algstr_0(X0) [cnf transformation 108] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
263. ~l6_algstr_0(X0) | ~v5_vectsp_1(X0) | v2_struct_0(X0) | v1_vectsp_1(X0) [cnf transformation 112] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
264. ~l6_algstr_0(X0) | ~v5_vectsp_1(X0) | v2_struct_0(X0) | v2_vectsp_1(X0) [cnf transformation 112] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
270. v9_struct_0(k4_struct_0(X0),X0) | ~l2_struct_0(X0) [cnf transformation 118] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
291. ~v9_struct_0(k5_struct_0(X0),X0) | ~l4_struct_0(X0) | v6_struct_0(X0) [cnf transformation 147] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
298. k4_struct_0(X0) = k6_algstr_0(X0,X1,k4_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l6_algstr_0(X0) | ~v1_vectsp_1(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 157] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
299. k4_struct_0(X0) = k6_algstr_0(X0,k4_struct_0(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l6_algstr_0(X0) | ~v2_vectsp_1(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 159] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
338. 1 <=> k4_struct_0(sK0) = sK1 [avatar definition]
340. k4_struct_0(sK0) = sK1 <- (1) [avatar component clause 338] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
342. 2 <=> k4_struct_0(sK0) = sK2 [avatar definition]
344. k4_struct_0(sK0) = sK2 <- (2) [avatar component clause 342] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
345. 1 | 2 [avatar split clause 240,342,338]
347. 3 <=> k5_struct_0(sK0) = k6_algstr_0(sK0,sK1,sK2) [avatar definition]
349. k5_struct_0(sK0) = k6_algstr_0(sK0,sK1,sK2) <- (3) [avatar component clause 347] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
350. 3 [avatar split clause 239,347]
352. 4 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
354. m1_subset_1(sK2,u1_struct_0(sK0)) <- (4) [avatar component clause 352] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
355. 4 [avatar split clause 238,352]
357. 5 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
359. m1_subset_1(sK1,u1_struct_0(sK0)) <- (5) [avatar component clause 357] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
360. 5 [avatar split clause 237,357]
362. 6 <=> l6_algstr_0(sK0) [avatar definition]
364. l6_algstr_0(sK0) <- (6) [avatar component clause 362] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
365. 6 [avatar split clause 236,362]
367. 7 <=> v4_rlvect_1(sK0) [avatar definition]
369. v4_rlvect_1(sK0) <- (7) [avatar component clause 367] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
370. 7 [avatar split clause 235,367]
372. 8 <=> v3_rlvect_1(sK0) [avatar definition]
374. v3_rlvect_1(sK0) <- (8) [avatar component clause 372] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
375. 8 [avatar split clause 234,372]
377. 9 <=> v5_vectsp_1(sK0) [avatar definition]
379. v5_vectsp_1(sK0) <- (9) [avatar component clause 377] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
380. 9 [avatar split clause 233,377]
382. 10 <=> v13_algstr_0(sK0) [avatar definition]
384. v13_algstr_0(sK0) <- (10) [avatar component clause 382] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
385. 10 [avatar split clause 232,382]
387. 11 <=> v6_struct_0(sK0) [avatar definition]
389. ~v6_struct_0(sK0) <- (~11) [avatar component clause 387] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
390. ~11 [avatar split clause 231,387]
392. 12 <=> v2_struct_0(sK0) [avatar definition]
394. ~v2_struct_0(sK0) <- (~12) [avatar component clause 392] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
395. ~12 [avatar split clause 230,392]
567. l2_algstr_0(sK0) <- (6) [resolution 258,364] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
570. 44 <=> l2_algstr_0(sK0) [avatar definition]
572. l2_algstr_0(sK0) <- (44) [avatar component clause 570] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
573. 44 | ~6 [avatar split clause 567,362,570]
580. l2_struct_0(sK0) <- (44) [resolution 572,248] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
587. 47 <=> l2_struct_0(sK0) [avatar definition]
589. l2_struct_0(sK0) <- (47) [avatar component clause 587] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
590. 47 | ~44 [avatar split clause 580,570,587]
591. l5_algstr_0(sK0) <- (6) [resolution 259,364] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
594. 48 <=> l5_algstr_0(sK0) [avatar definition]
596. l5_algstr_0(sK0) <- (48) [avatar component clause 594] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
597. 48 | ~6 [avatar split clause 591,362,594]
603. l4_struct_0(sK0) <- (48) [resolution 596,245] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
606. 50 <=> l4_struct_0(sK0) [avatar definition]
608. l4_struct_0(sK0) <- (50) [avatar component clause 606] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
609. 50 | ~48 [avatar split clause 603,594,606]
696. 64 <=> v9_struct_0(sK2,sK0) [avatar definition]
698. v9_struct_0(sK2,sK0) <- (64) [avatar component clause 696] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
703. 65 <=> v9_struct_0(sK1,sK0) [avatar definition]
705. v9_struct_0(sK1,sK0) <- (65) [avatar component clause 703] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
709. v9_struct_0(sK2,sK0) | ~l2_struct_0(sK0) <- (2) [superposition 270,344] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
715. v9_struct_0(sK2,sK0) <- (2, 47) [subsumption resolution 709,589] {a:1,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
716. 64 | ~2 | ~47 [avatar split clause 715,587,342,696]
899. ~v5_vectsp_1(sK0) | v2_struct_0(sK0) | v1_vectsp_1(sK0) <- (6) [resolution 263,364] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
901. v2_struct_0(sK0) | v1_vectsp_1(sK0) <- (6, 9) [subsumption resolution 899,379] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
902. v1_vectsp_1(sK0) <- (6, 9, ~12) [subsumption resolution 901,394] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
904. 95 <=> v1_vectsp_1(sK0) [avatar definition]
906. v1_vectsp_1(sK0) <- (95) [avatar component clause 904] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
907. 95 | ~6 | ~9 | 12 [avatar split clause 902,392,377,362,904]
921. ~v5_vectsp_1(sK0) | v2_struct_0(sK0) | v2_vectsp_1(sK0) <- (6) [resolution 264,364] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
923. v2_struct_0(sK0) | v2_vectsp_1(sK0) <- (6, 9) [subsumption resolution 921,379] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
924. v2_vectsp_1(sK0) <- (6, 9, ~12) [subsumption resolution 923,394] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
926. 99 <=> v2_vectsp_1(sK0) [avatar definition]
928. v2_vectsp_1(sK0) <- (99) [avatar component clause 926] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
929. 99 | ~6 | ~9 | 12 [avatar split clause 924,392,377,362,926]
1033. sK2 = k6_algstr_0(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l6_algstr_0(sK0) | ~v1_vectsp_1(sK0) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (2) [superposition 298,344] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1036. sK2 = k6_algstr_0(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v1_vectsp_1(sK0) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (2, 6) [subsumption resolution 1033,364] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1037. sK2 = k6_algstr_0(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (2, 6, 95) [subsumption resolution 1036,906] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1038. sK2 = k6_algstr_0(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (2, 6, 7, 95) [subsumption resolution 1037,369] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1039. sK2 = k6_algstr_0(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (2, 6, 7, 8, 95) [subsumption resolution 1038,374] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1040. sK2 = k6_algstr_0(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (2, 6, 7, 8, 10, 95) [subsumption resolution 1039,384] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1041. sK2 = k6_algstr_0(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (2, 6, 7, 8, 10, ~12, 95) [subsumption resolution 1040,394] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1076. k5_struct_0(sK0) = sK2 | ~m1_subset_1(sK1,u1_struct_0(sK0)) <- (2, 3, 6, 7, 8, 10, ~12, 95) [superposition 349,1041] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1082. 111 <=> k5_struct_0(sK0) = sK2 [avatar definition]
1084. k5_struct_0(sK0) = sK2 <- (111) [avatar component clause 1082] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1086. k5_struct_0(sK0) = sK2 <- (2, 3, 5, 6, 7, 8, 10, ~12, 95) [subsumption resolution 1076,359] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1087. 111 | ~2 | ~3 | ~5 | ~6 | ~7 | ~8 | ~10 | 12 | ~95 [avatar split clause 1086,904,392,382,372,367,362,357,347,342,1082]
1092. ~v9_struct_0(sK2,sK0) | ~l4_struct_0(sK0) | v6_struct_0(sK0) <- (111) [superposition 291,1084] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1099. ~l4_struct_0(sK0) | v6_struct_0(sK0) <- (64, 111) [subsumption resolution 1092,698] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1100. v6_struct_0(sK0) <- (50, 64, 111) [subsumption resolution 1099,608] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1101. $false <- (~11, 50, 64, 111) [subsumption resolution 1100,389] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1102. 11 | ~50 | ~64 | ~111 [avatar contradiction clause 1101]
1103. sK1 = k6_algstr_0(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l6_algstr_0(sK0) | ~v2_vectsp_1(sK0) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (1) [superposition 299,340] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1106. v9_struct_0(sK1,sK0) | ~l2_struct_0(sK0) <- (1) [superposition 270,340] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1107. sK1 = k6_algstr_0(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v2_vectsp_1(sK0) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (1, 6) [subsumption resolution 1103,364] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1108. sK1 = k6_algstr_0(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (1, 6, 99) [subsumption resolution 1107,928] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1109. sK1 = k6_algstr_0(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (1, 6, 7, 99) [subsumption resolution 1108,369] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1110. sK1 = k6_algstr_0(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (1, 6, 7, 8, 99) [subsumption resolution 1109,374] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1111. sK1 = k6_algstr_0(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (1, 6, 7, 8, 10, 99) [subsumption resolution 1110,384] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1112. sK1 = k6_algstr_0(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (1, 6, 7, 8, 10, ~12, 99) [subsumption resolution 1111,394] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1119. v9_struct_0(sK1,sK0) <- (1, 47) [subsumption resolution 1106,589] {a:1,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1120. 65 | ~1 | ~47 [avatar split clause 1119,587,338,703]
1132. k5_struct_0(sK0) = sK1 | ~m1_subset_1(sK2,u1_struct_0(sK0)) <- (1, 3, 6, 7, 8, 10, ~12, 99) [superposition 349,1112] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1138. 113 <=> k5_struct_0(sK0) = sK1 [avatar definition]
1140. k5_struct_0(sK0) = sK1 <- (113) [avatar component clause 1138] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1142. k5_struct_0(sK0) = sK1 <- (1, 3, 4, 6, 7, 8, 10, ~12, 99) [subsumption resolution 1132,354] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1143. 113 | ~1 | ~3 | ~4 | ~6 | ~7 | ~8 | ~10 | 12 | ~99 [avatar split clause 1142,926,392,382,372,367,362,352,347,338,1138]
1152. ~v9_struct_0(sK1,sK0) | ~l4_struct_0(sK0) | v6_struct_0(sK0) <- (113) [superposition 291,1140] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1159. ~l4_struct_0(sK0) | v6_struct_0(sK0) <- (65, 113) [subsumption resolution 1152,705] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1160. v6_struct_0(sK0) <- (50, 65, 113) [subsumption resolution 1159,608] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1161. $false <- (~11, 50, 65, 113) [subsumption resolution 1160,389] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1162. 11 | ~50 | ~65 | ~113 [avatar contradiction clause 1161]
1163. $false [avatar sat refutation 345,350,355,360,365,370,375,380,385,390,395,573,590,597,609,716,907,929,1087,1102,1120,1143,1162] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:817,thDist:-817}
% SZS output end Proof for vectsp_2__t5_vectsp_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 84
% >>> Preprocessing
% Introduced skolems: 22
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 107
% Generated clauses: 461
% Active clauses: 164
% Passive clauses: 217
% Final active clauses: 148
% Final passive clauses: 46
% Main loop iterations started: 169
% 
% >>> Simplifying Inferences
% Duplicate literals: 16
% Fw subsumption resolutions: 62
% Fw demodulations: 1
% Bw demodulations: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 37
% 
% >>> Generating Inferences
% Binary resolution: 124
% Unit resulting resolution: 1
% Forward superposition: 16
% Backward superposition: 22
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 23
% Split components: 48
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 124
% SAT solver unit clauses: 27
% SAT solver binary clauses: 47
% 
% Memory used [KB]: 5628
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% sat proof minimization: 0.001 s
% forward subsumption: 0.001 s
% other: 0.008 s ( own 0.004 s ) 

