i: [1,51,0,2,1,1,0,0] t133_pboole
a: [3,82,0,2,1,51]
s: 82
i: [1,49,0,2,1,1,0,0] t133_pboole
a: [3,92,0,2,1,49]
s: 92
i: [1,47,0,2,1,1,0,0] t133_pboole
a: [3,102,0,2,1,47]
s: 102
i: [1,45,0,2,1,1,0,0] t133_pboole
a: [3,112,0,2,1,45]
s: 112
i: [1,43,0,2,1,1,0,0] t133_pboole
a: [3,122,0,2,1,43]
s: 122
i: [1,41,0,2,1,1,0,0] t133_pboole
a: [3,132,0,2,1,41]
s: 132
i: [1,48,0,3,1,1,0,0] t133_pboole
a: [3,97,0,3,1,48]
s: 97
i: [1,46,0,3,1,1,0,0] t133_pboole
a: [3,107,0,3,1,46]
s: 107
i: [1,44,0,3,1,1,0,0] t133_pboole
a: [3,117,0,3,1,44]
s: 117
i: [1,42,0,3,1,1,0,0] t133_pboole
a: [3,127,0,3,1,42]
s: 127
i: [1,50,0,4,1,1,0,0] t133_pboole
a: [3,87,0,4,1,50]
s: 87
i: [1,76,0,2,1,0,0,1] rc3_funct_1
a: [3,137,0,2,1,76]
s: 137
i: [1,75,0,2,1,0,0,1] rc3_funct_1
a: [3,142,0,2,1,75]
s: 142
i: [1,74,0,2,1,0,0,1] rc3_funct_1
a: [3,147,0,2,1,74]
s: 147
i: [1,78,0,2,1,0,0,1] rc1_funct_1
a: [3,152,0,2,1,78]
s: 152
i: [1,77,0,2,1,0,0,1] rc1_funct_1
a: [3,157,0,2,1,77]
s: 157
i: [1,54,0,3,1,0,0,1] rc9_funct_1
s: 54
i: [1,55,0,3,1,0,0,1] rc9_funct_1
s: 55
i: [1,57,0,3,1,0,0,1] rc9_funct_1
s: 57
i: [1,58,0,3,1,0,0,1] rc3_pboole
s: 58
i: [1,59,0,3,1,0,0,1] rc3_pboole
s: 59
i: [1,61,0,3,1,0,0,1] rc3_pboole
s: 61
i: [1,63,0,3,1,0,0,1] rc4_funcop_1
s: 63
i: [1,64,0,3,1,0,0,1] rc4_funcop_1
s: 64
i: [1,66,0,3,1,0,0,1] rc4_funcop_1
s: 66
i: [1,68,0,3,1,0,0,1] rc1_pboole
s: 68
i: [1,70,0,3,1,0,0,1] rc1_pboole
s: 70
i: [1,56,0,4,1,0,0,1] rc9_funct_1
s: 56
i: [1,60,0,4,1,0,0,1] rc3_pboole
s: 60
i: [1,62,0,4,1,0,0,1] rc3_pboole
s: 62
i: [1,65,0,4,1,0,0,1] rc4_funcop_1
s: 65
i: [1,67,0,4,1,0,0,1] rc4_funcop_1
s: 67
i: [1,69,0,4,1,0,0,1] rc1_pboole
s: 69
i: [1,71,0,4,1,0,0,1] rc1_pboole
s: 71
i: [1,72,0,24,9,0,0,2] reflexivity_r2_pboole
s: 72
i: [1,73,0,24,9,0,0,1] reflexivity_r7_pboole
s: 73
i: [1,52,0,28,11,0,0,2] t131_pboole
s: 52
i: [1,53,0,30,11,0,0,1] t132_pboole
s: 53
d: [2,161,1,26,10,1,40,53,87]
d: [2,164,1,24,9,2,21,161,92]
d: [2,165,1,21,8,3,21,164,97]
d: [2,166,1,19,7,4,21,165,102]
d: [2,167,1,16,6,5,21,166,107]
d: [2,168,1,14,5,6,21,167,112]
d: [2,169,1,11,4,7,21,168,117]
d: [2,170,1,9,3,8,21,169,122]
d: [2,171,1,6,2,9,21,170,127]
d: [2,172,1,4,1,10,21,171,132]
a: [3,176,1,4,1,172]
s: 176
d: [2,178,2,0,0,11,52,132,102,122,82,112,92,127,117,107,97,176,52]
e: 178
d: [2,179,2,24,10,1,40,176,52]
d: [2,181,2,22,9,2,21,179,82]
d: [2,182,2,20,8,3,21,181,92]
d: [2,183,2,17,7,4,21,182,97]
d: [2,184,2,15,6,5,21,183,102]
d: [2,185,2,12,5,6,21,184,107]
d: [2,186,2,10,4,7,21,185,112]
d: [2,187,2,7,3,8,21,186,117]
d: [2,188,2,5,2,9,21,187,122]
d: [2,189,2,2,1,10,21,188,127]
d: [2,190,2,0,0,11,21,189,132]
e: 190
f: 190,172,41,42,43,44,45,46,47,48,49,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pboole__t133_pboole
% SZS output start Proof for pboole__t133_pboole
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r7_pboole(X0,X1,X2) & v2_relat_1(X1)) => v2_relat_1(X2)))) [input t133_pboole]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r7_pboole(X0,X1,X2) & v2_relat_1(X1)) => v2_relat_1(X2)))) [negated conjecture 1]
11. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X1,X2) & v2_relat_1(X1)) => v2_relat_1(X2)))) [input t131_pboole]
12. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r7_pboole(X0,X1,X2) & v2_relat_1(X1)) => r2_pboole(X0,X1,X2)))) [input t132_pboole]
13. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r7_pboole(X0,X1,X2) & v2_relat_1(X1)) => v2_relat_1(X2)))) [flattening 2]
14. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X1,X2) & v2_relat_1(X1)) => v2_relat_1(X2)))) [flattening 11]
15. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r7_pboole(X0,X1,X2) & v2_relat_1(X1)) => r2_pboole(X0,X1,X2)))) [flattening 12]
16. ? [X0,X1] : (? [X2] : ((~v2_relat_1(X2) & (r7_pboole(X0,X1,X2) & v2_relat_1(X1))) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 13]
17. ? [X0,X1] : (? [X2] : (~v2_relat_1(X2) & r7_pboole(X0,X1,X2) & v2_relat_1(X1) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 16]
18. ! [X0,X1] : (! [X2] : ((v2_relat_1(X2) | (~r2_pboole(X0,X1,X2) | ~v2_relat_1(X1))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 14]
19. ! [X0,X1] : (! [X2] : (v2_relat_1(X2) | ~r2_pboole(X0,X1,X2) | ~v2_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 18]
20. ! [X0,X1] : (! [X2] : ((r2_pboole(X0,X1,X2) | (~r7_pboole(X0,X1,X2) | ~v2_relat_1(X1))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 15]
21. ! [X0,X1] : (! [X2] : (r2_pboole(X0,X1,X2) | ~r7_pboole(X0,X1,X2) | ~v2_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 20]
26. ? [X0,X1] : (? [X2] : (~v2_relat_1(X2) & r7_pboole(X0,X1,X2) & v2_relat_1(X1) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (~v2_relat_1(X2) & r7_pboole(sK0,sK1,X2) & v2_relat_1(sK1) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
27. ? [X2] : (~v2_relat_1(X2) & r7_pboole(sK0,sK1,X2) & v2_relat_1(sK1) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (~v2_relat_1(sK2) & r7_pboole(sK0,sK1,sK2) & v2_relat_1(sK1) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
28. (~v2_relat_1(sK2) & r7_pboole(sK0,sK1,sK2) & v2_relat_1(sK1) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 17,27,26]
41. v1_relat_1(sK1) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. v4_relat_1(sK1,sK0) [cnf transformation 28] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
43. v1_funct_1(sK1) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
44. v1_partfun1(sK1,sK0) [cnf transformation 28] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
45. v1_relat_1(sK2) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
46. v4_relat_1(sK2,sK0) [cnf transformation 28] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
47. v1_funct_1(sK2) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. v1_partfun1(sK2,sK0) [cnf transformation 28] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. v2_relat_1(sK1) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. r7_pboole(sK0,sK1,sK2) [cnf transformation 28] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. ~v2_relat_1(sK2) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. ~r2_pboole(X0,X1,X2) | v2_relat_1(X2) | ~v2_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 19] {a:0,w:28,wCS:11200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
53. ~r7_pboole(X0,X1,X2) | r2_pboole(X0,X1,X2) | ~v2_relat_1(X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 21] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
80. 1 <=> v2_relat_1(sK2) [avatar definition]
82. ~v2_relat_1(sK2) <- (~1) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
83. ~1 [avatar split clause 51,80]
85. 2 <=> r7_pboole(sK0,sK1,sK2) [avatar definition]
87. r7_pboole(sK0,sK1,sK2) <- (2) [avatar component clause 85] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
88. 2 [avatar split clause 50,85]
90. 3 <=> v2_relat_1(sK1) [avatar definition]
92. v2_relat_1(sK1) <- (3) [avatar component clause 90] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
93. 3 [avatar split clause 49,90]
95. 4 <=> v1_partfun1(sK2,sK0) [avatar definition]
97. v1_partfun1(sK2,sK0) <- (4) [avatar component clause 95] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. 4 [avatar split clause 48,95]
100. 5 <=> v1_funct_1(sK2) [avatar definition]
102. v1_funct_1(sK2) <- (5) [avatar component clause 100] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
103. 5 [avatar split clause 47,100]
105. 6 <=> v4_relat_1(sK2,sK0) [avatar definition]
107. v4_relat_1(sK2,sK0) <- (6) [avatar component clause 105] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
108. 6 [avatar split clause 46,105]
110. 7 <=> v1_relat_1(sK2) [avatar definition]
112. v1_relat_1(sK2) <- (7) [avatar component clause 110] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. 7 [avatar split clause 45,110]
115. 8 <=> v1_partfun1(sK1,sK0) [avatar definition]
117. v1_partfun1(sK1,sK0) <- (8) [avatar component clause 115] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. 8 [avatar split clause 44,115]
120. 9 <=> v1_funct_1(sK1) [avatar definition]
122. v1_funct_1(sK1) <- (9) [avatar component clause 120] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
123. 9 [avatar split clause 43,120]
125. 10 <=> v4_relat_1(sK1,sK0) [avatar definition]
127. v4_relat_1(sK1,sK0) <- (10) [avatar component clause 125] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. 10 [avatar split clause 42,125]
130. 11 <=> v1_relat_1(sK1) [avatar definition]
132. v1_relat_1(sK1) <- (11) [avatar component clause 130] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
133. 11 [avatar split clause 41,130]
161. r2_pboole(sK0,sK1,sK2) | ~v2_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2) [resolution 53,87] {a:1,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
164. r2_pboole(sK0,sK1,sK2) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3) [subsumption resolution 161,92] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
165. r2_pboole(sK0,sK1,sK2) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4) [subsumption resolution 164,97] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
166. r2_pboole(sK0,sK1,sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5) [subsumption resolution 165,102] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
167. r2_pboole(sK0,sK1,sK2) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6) [subsumption resolution 166,107] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
168. r2_pboole(sK0,sK1,sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7) [subsumption resolution 167,112] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
169. r2_pboole(sK0,sK1,sK2) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8) [subsumption resolution 168,117] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
170. r2_pboole(sK0,sK1,sK2) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, 9) [subsumption resolution 169,122] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
171. r2_pboole(sK0,sK1,sK2) | ~v1_relat_1(sK1) <- (2, 3, 4, 5, 6, 7, 8, 9, 10) [subsumption resolution 170,127] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
172. r2_pboole(sK0,sK1,sK2) <- (2, 3, 4, 5, 6, 7, 8, 9, 10, 11) [subsumption resolution 171,132] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
174. 17 <=> r2_pboole(sK0,sK1,sK2) [avatar definition]
176. r2_pboole(sK0,sK1,sK2) <- (17) [avatar component clause 174] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
177. 17 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 [avatar split clause 172,130,125,120,115,110,105,100,95,90,85,174]
179. v2_relat_1(sK2) | ~v2_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (17) [resolution 176,52] {a:2,w:24,wCS:2400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
181. ~v2_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 17) [subsumption resolution 179,82] {a:2,w:22,wCS:2200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
182. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 3, 17) [subsumption resolution 181,92] {a:2,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
183. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 3, 4, 17) [subsumption resolution 182,97] {a:2,w:17,wCS:1700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
184. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 17) [subsumption resolution 183,102] {a:2,w:15,wCS:1500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
185. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 6, 17) [subsumption resolution 184,107] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
186. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 6, 7, 17) [subsumption resolution 185,112] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
187. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 6, 7, 8, 17) [subsumption resolution 186,117] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
188. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 6, 7, 8, 9, 17) [subsumption resolution 187,122] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
189. ~v1_relat_1(sK1) <- (~1, 3, 4, 5, 6, 7, 8, 9, 10, 17) [subsumption resolution 188,127] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
190. $false <- (~1, 3, 4, 5, 6, 7, 8, 9, 10, 11, 17) [subsumption resolution 189,132] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
191. 1 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~17 [avatar contradiction clause 190]
192. $false [avatar sat refutation 83,88,93,98,103,108,113,118,123,128,133,177,191] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:114,thDist:-114}
% SZS output end Proof for pboole__t133_pboole
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 38
% Generated clauses: 81
% Active clauses: 39
% Passive clauses: 39
% Final active clauses: 39
% Main loop iterations started: 39
% 
% >>> Simplifying Inferences
% Duplicate literals: 16
% Fw subsumption resolutions: 19
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 16
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit clause index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

