i: [1,300,0,2,1,1,0,0] t144_xreal_1
a: [3,477,0,2,1,300]
s: 477
i: [1,299,0,2,1,1,0,0] t144_xreal_1
a: [3,482,0,2,1,299]
s: 482
i: [1,303,0,6,2,1,0,0] t144_xreal_1
a: [3,458,0,3,1,303]
s: 458
i: [1,302,0,6,2,1,0,0] t144_xreal_1
a: [3,463,0,3,1,302]
s: 463
i: [1,301,0,5,1,1,0,0] t144_xreal_1
a: [3,472,0,5,1,301]
s: 472
i: [1,304,0,2,1,0,0,255] fc1_numbers
a: [3,487,0,2,1,304]
s: 487
i: [1,305,0,2,1,0,0,255] spc1_boole
a: [3,492,0,2,1,305]
s: 492
i: [1,306,0,2,1,0,0,255] fc1_xboole_0
a: [3,497,0,2,1,306]
s: 497
i: [1,316,0,2,1,0,0,255] fc6_ordinal1
a: [3,542,0,2,1,316]
s: 542
i: [1,315,0,2,1,0,0,255] fc6_ordinal1
a: [3,547,0,2,1,315]
s: 547
i: [1,317,0,2,1,0,0,255] spc1_numerals
a: [3,567,0,2,1,317]
s: 567
i: [1,408,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,571,0,2,1,408]
s: 571
i: [1,418,0,2,1,0,0,2] rc2_xboole_0
a: [3,579,0,2,1,418]
s: 579
i: [1,419,0,2,1,0,0,255] rc4_ordinal1
a: [3,584,0,2,1,419]
s: 584
i: [1,420,0,2,1,0,0,255] rc1_ordinal1
a: [3,589,0,2,1,420]
s: 589
i: [1,421,0,2,1,0,0,1] rc1_xreal_0
a: [3,594,0,2,1,421]
s: 594
i: [1,422,0,2,1,0,0,2] rc1_xxreal_0
a: [3,599,0,2,1,422]
s: 599
d: [2,741,0,0,0,2,21,599,571]
e: 741
r: 571
s: 599
i: [1,423,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,604,0,2,1,423]
s: 604
i: [1,424,0,2,1,0,0,2] rc1_xboole_0
a: [3,609,0,2,1,424]
s: 609
i: [1,428,0,2,1,0,0,255] rc2_ordinal1
a: [3,614,0,2,1,428]
s: 614
i: [1,427,0,2,1,0,0,255] rc2_ordinal1
a: [3,619,0,2,1,427]
s: 619
i: [1,426,0,2,1,0,0,255] rc2_ordinal1
a: [3,624,0,2,1,426]
s: 624
i: [1,425,0,2,1,0,0,255] rc2_ordinal1
a: [3,629,0,2,1,425]
s: 629
i: [1,430,0,2,1,0,0,255] rc5_ordinal1
a: [3,634,0,2,1,430]
s: 634
i: [1,429,0,2,1,0,0,255] rc5_ordinal1
a: [3,639,0,2,1,429]
s: 639
i: [1,432,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,644,0,2,1,432]
s: 644
i: [1,431,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,649,0,2,1,431]
s: 649
i: [1,434,0,2,1,0,0,2] rc3_xxreal_0
a: [3,654,0,2,1,434]
s: 654
i: [1,433,0,2,1,0,0,2] rc3_xxreal_0
a: [3,659,0,2,1,433]
s: 659
i: [1,436,0,2,1,0,0,2] rc2_xxreal_0
a: [3,664,0,2,1,436]
s: 664
i: [1,435,0,2,1,0,0,2] rc2_xxreal_0
a: [3,669,0,2,1,435]
s: 669
i: [1,440,0,2,1,0,0,2] rc3_xreal_0
a: [3,674,0,2,1,440]
s: 674
i: [1,439,0,2,1,0,0,2] rc3_xreal_0
a: [3,679,0,2,1,439]
s: 679
i: [1,438,0,2,1,0,0,2] rc3_xreal_0
a: [3,684,0,2,1,438]
s: 684
i: [1,437,0,2,1,0,0,2] rc3_xreal_0
a: [3,689,0,2,1,437]
s: 689
i: [1,444,0,2,1,0,0,2] rc2_xreal_0
a: [3,694,0,2,1,444]
s: 694
i: [1,443,0,2,1,0,0,2] rc2_xreal_0
a: [3,699,0,2,1,443]
s: 699
i: [1,442,0,2,1,0,0,2] rc2_xreal_0
a: [3,704,0,2,1,442]
s: 704
i: [1,441,0,2,1,0,0,2] rc2_xreal_0
a: [3,709,0,2,1,441]
s: 709
i: [1,446,0,2,1,0,0,2] rc4_xxreal_0
a: [3,714,0,2,1,446]
s: 714
i: [1,445,0,2,1,0,0,2] rc4_xxreal_0
a: [3,719,0,2,1,445]
s: 719
i: [1,450,0,2,1,0,0,2] rc4_xreal_0
a: [3,724,0,2,1,450]
s: 724
i: [1,449,0,2,1,0,0,2] rc4_xreal_0
a: [3,729,0,2,1,449]
s: 729
i: [1,448,0,2,1,0,0,2] rc4_xreal_0
a: [3,734,0,2,1,448]
s: 734
i: [1,447,0,2,1,0,0,2] rc4_xreal_0
a: [3,739,0,2,1,447]
s: 739
i: [1,390,0,3,1,0,0,255] reflexivity_r1_tarski
s: 390
i: [1,307,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,502,0,3,1,307]
s: 502
i: [1,308,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,507,0,3,1,308]
s: 507
r: 497
d: [2,743,0,2,1,2,23,497,507]
a: [3,747,0,2,1,743]
s: 747
i: [1,309,0,3,1,0,0,255] t1_numerals
a: [3,512,0,3,1,309]
s: 512
i: [1,310,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,517,0,3,1,310]
s: 517
i: [1,320,0,3,1,0,0,255] spc1_numerals
a: [3,552,0,3,1,320]
s: 552
i: [1,319,0,3,1,0,0,255] spc1_numerals
a: [3,557,0,3,1,319]
s: 557
d: [2,749,0,3,1,2,22,512,507]
a: [3,753,0,3,1,749]
s: 753
d: [2,755,0,3,1,2,22,557,502]
a: [3,759,0,3,1,755]
s: 759
i: [1,322,0,4,2,0,0,255] cc6_ordinal1
s: 322
i: [1,323,0,4,2,0,0,255] cc2_xreal_0
s: 323
d: [2,761,1,2,1,1,40,322,584]
a: [3,766,1,2,1,761]
s: 766
i: [1,324,0,4,2,0,0,255] cc2_xxreal_0
s: 324
i: [1,325,0,4,2,0,0,255] cc2_xcmplx_0
s: 325
d: [2,762,1,2,1,1,40,322,634]
a: [3,771,1,2,1,762]
s: 771
i: [1,326,0,4,2,0,0,255] cc1_ordinal1
s: 326
d: [2,773,1,2,1,1,40,323,584]
a: [3,778,1,2,1,773]
s: 778
i: [1,327,0,4,2,0,0,255] cc1_ordinal1
s: 327
i: [1,329,0,4,2,0,0,2] cc4_xreal_0
s: 329
d: [2,859,1,2,1,1,40,329,482]
a: [3,869,1,2,1,859]
s: 869
i: [1,330,0,4,2,0,0,1] cc3_xreal_0
s: 330
d: [2,860,1,2,1,1,40,329,477]
a: [3,874,1,2,1,860]
s: 874
i: [1,349,0,4,2,0,0,255] cc9_ordinal1
s: 349
i: [1,350,0,4,2,0,0,255] cc7_ordinal1
s: 350
d: [2,882,1,2,1,1,40,330,482]
a: [3,892,1,2,1,882]
s: 892
i: [1,351,0,4,2,0,0,255] cc3_ordinal1
s: 351
d: [2,883,1,2,1,1,40,330,477]
a: [3,897,1,2,1,883]
s: 897
i: [1,389,0,4,1,0,0,255] existence_m1_subset_1
s: 389
i: [1,311,0,4,1,0,0,255] dt_k5_numbers
a: [3,522,0,4,1,311]
s: 522
d: [2,774,1,2,1,1,40,323,634]
a: [3,783,1,2,1,774]
s: 783
i: [1,312,0,4,1,0,0,255] dt_k6_numbers
a: [3,527,0,4,1,312]
s: 527
d: [2,785,1,2,1,1,40,324,584]
a: [3,790,1,2,1,785]
s: 790
i: [1,318,0,4,1,0,0,255] spc1_numerals
a: [3,562,0,4,1,318]
s: 562
d: [2,977,0,4,1,2,22,522,502]
a: [3,981,0,4,1,977]
s: 981
d: [2,786,1,2,1,1,40,324,634]
a: [3,795,1,2,1,786]
s: 795
d: [2,987,0,4,1,2,22,527,502]
a: [3,991,0,4,1,987]
s: 991
d: [2,797,1,2,1,1,40,325,584]
a: [3,802,1,2,1,797]
s: 802
d: [2,993,0,4,1,2,22,562,502]
a: [3,997,0,4,1,993]
s: 997
i: [1,331,0,5,2,0,0,255] fc4_xreal_0
s: 331
d: [2,798,1,2,1,1,40,325,634]
a: [3,807,1,2,1,798]
s: 807
i: [1,332,0,5,2,0,0,255] fc4_xreal_0
s: 332
d: [2,809,1,2,1,1,40,326,542]
a: [3,817,1,2,1,809]
s: 817
i: [1,341,0,5,2,0,0,255] dt_k5_xcmplx_0
s: 341
i: [1,352,0,5,2,0,0,255] t6_boole
s: 352
d: [2,811,1,2,1,1,40,326,589]
a: [3,827,1,2,1,811]
s: 827
i: [1,353,0,5,2,0,0,255] cc8_ordinal1
s: 353
d: [2,834,1,2,1,1,40,327,542]
a: [3,842,1,2,1,834]
s: 842
i: [1,354,0,5,2,0,0,255] cc1_xreal_0
s: 354
i: [1,355,0,5,2,0,0,255] cc1_xcmplx_0
s: 355
d: [2,836,1,2,1,1,40,327,589]
a: [3,852,1,2,1,836]
s: 852
i: [1,386,0,5,2,0,0,255] t7_boole
s: 386
d: [2,862,1,2,1,1,40,329,594]
a: [3,880,1,2,1,862]
s: 880
i: [1,313,0,5,1,0,0,255] rqRealDiv__k7_xcmplx_0__r1_r1_r1
a: [3,532,0,5,1,313]
s: 532
i: [1,314,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,537,0,5,1,314]
s: 537
d: [2,885,1,2,1,1,40,330,594]
a: [3,903,1,2,1,885]
s: 903
a: [3,574,0,5,2,408]
s: 574
d: [2,905,1,2,1,1,40,349,747]
a: [3,912,1,2,1,905]
s: 912
d: [2,1001,0,5,2,1,22,352,507]
s: 1001
i: [1,366,0,6,3,0,0,255] cc2_ordinal1
s: 366
d: [2,1041,1,5,2,2,44,458,1001]
s: 1041
i: [1,369,0,6,3,0,0,2] cc5_xxreal_0
s: 369
d: [2,1042,1,5,2,2,44,463,1001]
s: 1042
i: [1,371,0,6,3,0,0,2] cc5_xxreal_0
s: 371
i: [1,372,0,6,3,0,0,2] cc3_xxreal_0
s: 372
d: [2,1043,1,7,2,2,44,472,1001]
s: 1043
i: [1,374,0,6,3,0,0,2] cc3_xxreal_0
s: 374
d: [2,1094,2,8,2,3,40,1043,574]
a: [3,1103,2,4,1,1094]
s: 1103
i: [1,378,0,6,3,0,0,2] cc7_xxreal_0
s: 378
i: [1,379,0,6,3,0,0,2] cc7_xxreal_0
s: 379
d: [2,906,1,2,1,1,40,349,609]
a: [3,917,1,2,1,906]
s: 917
i: [1,382,0,6,2,0,0,255] t1_subset
s: 382
d: [2,907,1,2,1,1,40,349,719]
a: [3,922,1,2,1,907]
s: 922
i: [1,391,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 391
i: [1,328,0,7,3,0,0,255] cc5_ordinal1
s: 328
d: [2,908,1,2,1,1,40,349,739]
a: [3,927,1,2,1,908]
s: 927
i: [1,342,0,7,2,0,0,1] t5_arithm
s: 342
d: [2,929,1,2,1,1,40,350,747]
a: [3,936,1,2,1,929]
s: 936
i: [1,343,0,7,2,0,0,1] t2_arithm
s: 343
i: [1,344,0,7,2,0,0,255] involutiveness_k5_xcmplx_0
s: 344
d: [2,930,1,2,1,1,40,350,609]
a: [3,941,1,2,1,930]
s: 941
i: [1,345,0,7,2,0,0,255] t6_arithm
s: 345
d: [2,931,1,2,1,1,40,350,719]
a: [3,946,1,2,1,931]
s: 946
i: [1,346,0,7,2,0,0,255] t3_arithm
s: 346
i: [1,356,0,7,3,0,0,255] fc30_xreal_0
s: 356
d: [2,932,1,2,1,1,40,350,739]
a: [3,951,1,2,1,932]
s: 951
i: [1,357,0,7,3,0,0,255] fc30_xreal_0
s: 357
d: [2,953,1,2,1,1,40,351,747]
a: [3,960,1,2,1,953]
s: 960
i: [1,362,0,7,3,0,0,255] fc7_xcmplx_0
s: 362
i: [1,363,0,7,3,0,0,255] fc7_xcmplx_0
s: 363
d: [2,954,1,2,1,1,40,351,609]
a: [3,965,1,2,1,954]
s: 965
i: [1,364,0,7,3,0,0,255] fc28_xreal_0
s: 364
d: [2,955,1,2,1,1,40,351,719]
a: [3,970,1,2,1,955]
s: 970
i: [1,365,0,7,3,0,0,255] fc28_xreal_0
s: 365
i: [1,375,0,7,3,0,0,255] fc29_xreal_0
s: 375
d: [2,956,1,2,1,1,40,351,739]
a: [3,975,1,2,1,956]
s: 975
i: [1,376,0,7,3,0,0,255] fc29_xreal_0
s: 376
d: [2,1003,1,2,1,1,40,353,759]
a: [3,1009,1,2,1,1003]
s: 1009
i: [1,380,0,7,3,0,0,255] fc27_xreal_0
s: 380
i: [1,381,0,7,3,0,0,255] fc27_xreal_0
s: 381
d: [2,1016,1,2,1,1,40,354,552]
a: [3,1021,1,2,1,1016]
s: 1021
i: [1,384,0,7,2,0,0,255] t3_subset
s: 384
d: [2,1028,1,2,1,1,40,355,552]
a: [3,1033,1,2,1,1028]
s: 1033
i: [1,385,0,7,3,0,0,2] t8_boole
s: 385
i: [1,321,0,8,3,0,0,255] cc10_ordinal1
s: 321
d: [2,1051,1,4,2,1,40,369,654]
d: [2,1053,1,2,1,2,21,1051,659]
a: [3,1057,1,2,1,1053]
s: 1057
i: [1,347,0,8,2,0,0,255] spc3_arithm
s: 347
d: [2,1052,1,4,2,1,40,369,679]
d: [2,1059,1,2,1,2,21,1052,684]
a: [3,1063,1,2,1,1059]
s: 1063
i: [1,359,0,8,4,0,0,2] cc4_xxreal_0
s: 359
i: [1,361,0,8,4,0,0,2] cc6_xxreal_0
s: 361
d: [2,1095,2,6,2,3,40,1043,1041]
a: [3,1107,2,2,1,1095]
s: 1107
r: 1107
i: [1,367,0,8,4,0,0,2] cc8_xxreal_0
s: 367
a: [3,1108,2,2,1,1095]
s: 1108
i: [1,383,0,8,3,0,0,255] t2_subset
s: 383
i: [1,406,0,8,3,0,0,1] fc8_xreal_0
s: 406
d: [2,1303,2,6,3,1,40,359,869]
a: [3,1318,2,2,1,1303]
s: 1318
i: [1,407,0,8,3,0,0,2] fc6_xreal_0
s: 407
d: [2,1304,2,6,3,1,40,359,874]
a: [3,1327,2,2,1,1304]
s: 1327
r: 1327
i: [1,410,0,8,3,0,0,1] fc5_xcmplx_0
s: 410
i: [1,411,0,8,3,0,0,2] fc3_xcmplx_0
s: 411
a: [3,1326,2,2,1,1304]
s: 1326
i: [1,388,0,9,3,0,0,255] t5_subset
s: 388
a: [3,1331,2,2,1,1304]
s: 1331
i: [1,387,0,10,3,0,0,255] t4_subset
s: 387
i: [1,409,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 409
a: [3,1099,2,4,1,1094]
s: 1099
i: [1,335,0,11,5,0,0,1] t2_real
s: 335
a: [3,1321,2,2,1,1303]
s: 1321
i: [1,336,0,11,5,0,0,1] t1_real
s: 336
r: 1331
i: [1,337,0,11,5,0,0,1] t8_real
s: 337
a: [3,1330,2,2,1,1304]
s: 1330
i: [1,338,0,11,5,0,0,1] t7_real
s: 338
a: [3,1335,2,2,1,1304]
s: 1335
i: [1,339,0,11,5,0,0,1] t3_real
s: 339
i: [1,340,0,11,5,0,0,1] t4_real
s: 340
d: [2,1559,1,12,4,1,40,337,472]
a: [3,1568,1,4,1,1559]
s: 1568
d: [2,1577,2,0,0,3,52,482,477,1568,406]
e: 1577
d: [2,1578,2,4,2,1,40,1568,406]
d: [2,1580,2,2,1,2,21,1578,477]
d: [2,1581,2,0,0,3,21,1580,482]
e: 1581
r: 1568
i: [1,412,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 412
d: [2,1506,1,8,4,1,40,335,458]
d: [2,1519,1,6,3,2,21,1506,482]
a: [3,1523,1,2,1,1519]
s: 1523
r: 1523
i: [1,348,0,12,3,0,0,255] d9_xcmplx_0
s: 348
i: [1,392,0,12,5,0,0,1] fc33_xreal_0
s: 392
a: [3,1522,1,2,1,1519]
s: 1522
i: [1,393,0,12,5,0,0,2] fc26_xreal_0
s: 393
a: [3,1567,1,4,1,1559]
s: 1567
d: [2,1656,3,0,0,2,52,1099,1567,329]
e: 1656
d: [2,1658,2,4,1,1,40,1567,329]
d: [2,1660,2,0,0,2,21,1658,1099]
e: 1660
r: 1099
i: [1,394,0,12,5,0,0,1] fc31_xreal_0
s: 394
i: [1,395,0,12,5,0,0,1] fc32_xreal_0
s: 395
d: [2,1537,1,8,4,1,40,336,463]
d: [2,1543,1,6,3,2,21,1537,1331]
d: [2,1544,1,4,2,3,21,1543,477]
a: [3,1547,1,2,1,1544]
s: 1547
i: [1,400,0,12,5,0,0,1] fc9_xcmplx_0
s: 400
a: [3,1572,1,4,1,1559]
s: 1572
d: [2,1693,3,0,0,5,52,1321,482,1330,477,1572,394]
e: 1693
d: [2,1695,2,8,4,1,40,1572,394]
d: [2,1698,2,6,3,2,21,1695,477]
d: [2,1699,2,4,2,3,21,1698,1330]
d: [2,1700,2,2,1,4,21,1699,482]
d: [2,1701,2,0,0,5,21,1700,1321]
e: 1701
r: 1041
r: 458
r: 1321
a: [3,454,0,3,1,303]
s: 454
r: 1326
r: 1335
a: [3,457,0,3,1,303]
s: 457
r: 457
r: 1330
r: 1547
a: [3,1548,1,2,1,1544]
s: 1548
r: 1042
r: 463
r: 1548
s: 458
s: 1547
a: [3,462,0,3,1,302]
s: 462
a: [3,467,0,3,1,302]
s: 467
r: 458
r: 1318
r: 454
r: 1108
d: [2,1002,1,2,1,1,40,353,753]
d: [2,1130,2,2,1,1,40,936,324]
d: [2,1131,2,2,1,1,40,936,323]
d: [2,1557,1,6,3,2,21,1537,477]
d: [2,1617,2,4,2,1,40,1523,369]
d: [2,1620,2,2,1,2,21,1617,747]
d: [2,1705,1,8,4,1,40,454,335]
d: [2,1711,1,6,3,2,21,1705,1522]
d: [2,1712,1,4,2,3,21,1711,477]
d: [2,1715,1,7,3,1,40,457,409]
d: [2,1720,1,5,2,2,21,1715,869]
d: [2,1722,2,4,2,3,21,1698,482]
d: [2,1726,2,4,2,1,40,1548,372]
d: [2,1728,2,2,1,2,21,1726,747]
d: [2,1731,1,8,4,1,40,458,335]
d: [2,1733,1,6,3,2,21,1731,1522]
d: [2,1734,1,4,2,3,21,1733,482]
d: [2,1743,1,8,4,1,40,467,336]
d: [2,1750,1,6,3,2,21,1743,1547]
d: [2,1751,1,4,2,3,21,1750,482]
d: [2,1694,2,8,4,1,40,1572,395]
d: [2,1753,2,6,3,2,21,1694,482]
d: [2,1754,2,4,2,3,21,1753,477]
d: [2,1738,1,7,3,1,40,462,409]
d: [2,1756,1,5,2,2,21,1738,874]
f: 1756,1754,1751,1734,1728,1722,1720,1712,1620,1581,1559,1557,1131,1130,1002,860,859,749,743,309,308,306,299,300,301,302,303
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t144_xreal_1
% SZS output start Proof for xreal_1__t144_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~(~r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k6_numbers,X1)) & ~(~r1_xxreal_0(X0,k6_numbers) & ~r1_xxreal_0(X1,k6_numbers)) & ~r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers)))) [input t144_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~(~r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k6_numbers,X1)) & ~(~r1_xxreal_0(X0,k6_numbers) & ~r1_xxreal_0(X1,k6_numbers)) & ~r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers)))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
14. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
19. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
28. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
44. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
53. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v2_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc31_xreal_0]
54. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v2_xxreal_0(k7_xcmplx_0(X1,X0))) [input fc32_xreal_0]
65. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
85. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
100. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
101. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
104. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
120. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
132. ? [X0] : (? [X1] : (((r1_xxreal_0(k6_numbers,X0) | r1_xxreal_0(k6_numbers,X1)) & (r1_xxreal_0(X0,k6_numbers) | r1_xxreal_0(X1,k6_numbers)) & ~r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
133. ? [X0] : (? [X1] : ((r1_xxreal_0(k6_numbers,X0) | r1_xxreal_0(k6_numbers,X1)) & (r1_xxreal_0(X0,k6_numbers) | r1_xxreal_0(X1,k6_numbers)) & ~r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 132]
136. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
137. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
141. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
148. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 104]
149. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 148]
150. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 101]
151. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 150]
152. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 120]
153. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 152]
172. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
189. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 19]
190. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 189]
191. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 14]
192. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 191]
213. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 53]
214. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 213]
215. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X1,X0)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 54]
216. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 215]
235. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 65]
236. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 235]
241. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 28]
242. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 241]
259. ? [X0] : (? [X1] : ((r1_xxreal_0(k6_numbers,X0) | r1_xxreal_0(k6_numbers,X1)) & (r1_xxreal_0(X0,k6_numbers) | r1_xxreal_0(X1,k6_numbers)) & ~r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : ((r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(k6_numbers,X1)) & (r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(X1,k6_numbers)) & ~r1_xxreal_0(k7_xcmplx_0(sK0,X1),k6_numbers) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
260. ? [X1] : ((r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(k6_numbers,X1)) & (r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(X1,k6_numbers)) & ~r1_xxreal_0(k7_xcmplx_0(sK0,X1),k6_numbers) & v1_xreal_0(X1)) => ((r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(k6_numbers,sK1)) & (r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(sK1,k6_numbers)) & ~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) & v1_xreal_0(sK1)) [choice axiom]
261. ((r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(k6_numbers,sK1)) & (r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(sK1,k6_numbers)) & ~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 133,260,259]
299. v1_xreal_0(sK0) [cnf transformation 261] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
300. v1_xreal_0(sK1) [cnf transformation 261] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
301. ~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) [cnf transformation 261] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
302. r1_xxreal_0(sK0,k6_numbers) | r1_xxreal_0(sK1,k6_numbers) [cnf transformation 261] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
303. r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(k6_numbers,sK1) [cnf transformation 261] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
306. v1_xboole_0(k1_xboole_0) [cnf transformation 44] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
308. k6_numbers = k1_xboole_0 [cnf transformation 85] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
309. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 100] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
323. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 136] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
324. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 137] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
329. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 141] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
335. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 149] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
336. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 151] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
337. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 153] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
353. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 172] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
369. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 190] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
372. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 192] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
394. ~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 214] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
395. ~v2_xxreal_0(k7_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 216] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
406. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 236] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
409. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 242] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
452. 1 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
454. r1_xxreal_0(k6_numbers,sK1) <- (1) [avatar component clause 452] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
456. 2 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
457. ~r1_xxreal_0(k6_numbers,sK0) <- (~2) [avatar component clause 456] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
458. r1_xxreal_0(k6_numbers,sK0) <- (2) [avatar component clause 456] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
459. 1 | 2 [avatar split clause 303,456,452]
461. 3 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
462. ~r1_xxreal_0(sK1,k6_numbers) <- (~3) [avatar component clause 461] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
463. r1_xxreal_0(sK1,k6_numbers) <- (3) [avatar component clause 461] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
465. 4 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
467. r1_xxreal_0(sK0,k6_numbers) <- (4) [avatar component clause 465] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
468. 3 | 4 [avatar split clause 302,465,461]
470. 5 <=> r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) [avatar definition]
472. ~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) <- (~5) [avatar component clause 470] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
473. ~5 [avatar split clause 301,470]
475. 6 <=> v1_xreal_0(sK1) [avatar definition]
477. v1_xreal_0(sK1) <- (6) [avatar component clause 475] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
478. 6 [avatar split clause 300,475]
480. 7 <=> v1_xreal_0(sK0) [avatar definition]
482. v1_xreal_0(sK0) <- (7) [avatar component clause 480] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
483. 7 [avatar split clause 299,480]
495. 10 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
497. v1_xboole_0(k1_xboole_0) <- (10) [avatar component clause 495] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
498. 10 [avatar split clause 306,495]
505. 12 <=> k6_numbers = k1_xboole_0 [avatar definition]
507. k6_numbers = k1_xboole_0 <- (12) [avatar component clause 505] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
508. 12 [avatar split clause 308,505]
510. 13 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
512. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (13) [avatar component clause 510] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
513. 13 [avatar split clause 309,510]
743. v1_xboole_0(k6_numbers) <- (10, 12) [backward demodulation 497,507] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
745. 60 <=> v1_xboole_0(k6_numbers) [avatar definition]
747. v1_xboole_0(k6_numbers) <- (60) [avatar component clause 745] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
748. 60 | ~10 | ~12 [avatar split clause 743,505,495,745]
749. m1_subset_1(k6_numbers,k4_ordinal1) <- (12, 13) [forward demodulation 512,507] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
751. 61 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
753. m1_subset_1(k6_numbers,k4_ordinal1) <- (61) [avatar component clause 751] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
754. 61 | ~12 | ~13 [avatar split clause 749,510,505,751]
859. v1_xxreal_0(sK0) <- (7) [resolution 329,482] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
860. v1_xxreal_0(sK1) <- (6) [resolution 329,477] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
867. 79 <=> v1_xxreal_0(sK0) [avatar definition]
869. v1_xxreal_0(sK0) <- (79) [avatar component clause 867] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
870. 79 | ~7 [avatar split clause 859,480,867]
872. 80 <=> v1_xxreal_0(sK1) [avatar definition]
874. v1_xxreal_0(sK1) <- (80) [avatar component clause 872] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
875. 80 | ~6 [avatar split clause 860,475,872]
934. 89 <=> v7_ordinal1(k6_numbers) [avatar definition]
936. v7_ordinal1(k6_numbers) <- (89) [avatar component clause 934] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1002. v7_ordinal1(k6_numbers) <- (61) [resolution 353,753] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1005. 89 | ~61 [avatar split clause 1002,751,934]
1130. v1_xxreal_0(k6_numbers) <- (89) [resolution 936,324] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1131. v1_xreal_0(k6_numbers) <- (89) [resolution 936,323] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1140. 118 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1143. 118 | ~89 [avatar split clause 1130,934,1140]
1145. 119 <=> v1_xreal_0(k6_numbers) [avatar definition]
1148. 119 | ~89 [avatar split clause 1131,934,1145]
1316. 137 <=> v2_xxreal_0(sK0) [avatar definition]
1320. 138 <=> v3_xxreal_0(sK0) [avatar definition]
1329. 140 <=> v2_xxreal_0(sK1) [avatar definition]
1333. 141 <=> v3_xxreal_0(sK1) [avatar definition]
1521. 167 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1522. ~v3_xxreal_0(k6_numbers) <- (~167) [avatar component clause 1521] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1523. v3_xxreal_0(k6_numbers) <- (167) [avatar component clause 1521] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1537. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (3) [resolution 336,463] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1546. 169 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1547. ~v2_xxreal_0(k6_numbers) <- (~169) [avatar component clause 1546] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1548. v2_xxreal_0(k6_numbers) <- (169) [avatar component clause 1546] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1557. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 6) [subsumption resolution 1537,477] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1558. ~119 | 169 | ~140 | ~3 | ~6 [avatar split clause 1557,475,461,1329,1546,1145]
1559. v3_xxreal_0(k6_numbers) | v2_xxreal_0(k7_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~5) [resolution 337,472] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1566. 171 <=> v1_xreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1568. ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~171) [avatar component clause 1566] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1570. 172 <=> v2_xxreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1572. v2_xxreal_0(k7_xcmplx_0(sK0,sK1)) <- (172) [avatar component clause 1570] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1573. ~171 | ~119 | 172 | 167 | 5 [avatar split clause 1559,470,1521,1570,1145,1566]
1578. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~171) [resolution 1568,406] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1580. ~v1_xreal_0(sK0) <- (6, ~171) [subsumption resolution 1578,477] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1581. $false <- (6, 7, ~171) [subsumption resolution 1580,482] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1582. ~6 | ~7 | 171 [avatar contradiction clause 1581]
1617. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (167) [resolution 1523,369] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1620. ~v1_xxreal_0(k6_numbers) <- (60, 167) [subsumption resolution 1617,747] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1621. ~118 | ~60 | ~167 [avatar split clause 1620,1521,745,1140]
1694. ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) | ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) <- (172) [resolution 1572,395] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1695. ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (172) [resolution 1572,394] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1698. v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (6, 172) [subsumption resolution 1695,477] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1705. ~v3_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (1) [resolution 454,335] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1711. ~v3_xxreal_0(sK1) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (1, ~167) [subsumption resolution 1705,1522] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1712. ~v3_xxreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (1, 6, ~167) [subsumption resolution 1711,477] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1713. ~119 | ~141 | ~1 | ~6 | 167 [avatar split clause 1712,1521,475,452,1333,1145]
1715. r1_xxreal_0(sK0,k6_numbers) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(sK0) <- (~2) [resolution 457,409] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1720. r1_xxreal_0(sK0,k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (~2, 79) [subsumption resolution 1715,869] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1721. ~118 | 4 | 2 | ~79 [avatar split clause 1720,867,456,465,1140]
1722. v2_xxreal_0(sK1) | v3_xxreal_0(sK0) <- (6, 7, 172) [subsumption resolution 1698,482] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1723. 138 | 140 | ~6 | ~7 | ~172 [avatar split clause 1722,1570,480,475,1329,1320]
1726. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (169) [resolution 1548,372] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1728. ~v1_xxreal_0(k6_numbers) <- (60, 169) [subsumption resolution 1726,747] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1729. ~118 | ~60 | ~169 [avatar split clause 1728,1546,745,1140]
1731. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (2) [resolution 458,335] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1733. ~v3_xxreal_0(sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (2, ~167) [subsumption resolution 1731,1522] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1734. ~v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (2, 7, ~167) [subsumption resolution 1733,482] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1735. ~119 | ~138 | ~2 | ~7 | 167 [avatar split clause 1734,1521,480,456,1320,1145]
1738. r1_xxreal_0(k6_numbers,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k6_numbers) <- (~3) [resolution 462,409] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1743. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (4) [resolution 467,336] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1750. ~v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (4, ~169) [subsumption resolution 1743,1547] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1751. ~v2_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (4, 7, ~169) [subsumption resolution 1750,482] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1752. ~119 | ~137 | ~4 | ~7 | 169 [avatar split clause 1751,1546,480,465,1316,1145]
1753. v2_xxreal_0(sK0) | ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) <- (7, 172) [subsumption resolution 1694,482] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1754. v2_xxreal_0(sK0) | v3_xxreal_0(sK1) <- (6, 7, 172) [subsumption resolution 1753,477] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1755. 141 | 137 | ~6 | ~7 | ~172 [avatar split clause 1754,1570,480,475,1316,1333]
1756. r1_xxreal_0(k6_numbers,sK1) | ~v1_xxreal_0(k6_numbers) <- (~3, 80) [subsumption resolution 1738,874] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1757. ~118 | 1 | 3 | ~80 [avatar split clause 1756,872,461,452,1140]
1758. $false [avatar sat refutation 459,468,473,478,483,498,508,513,748,754,870,875,1005,1143,1148,1558,1573,1582,1621,1713,1721,1723,1729,1735,1752,1755,1757] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1113,thDist:-1113}
% SZS output end Proof for xreal_1__t144_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 119
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 152
% Generated clauses: 762
% Active clauses: 181
% Passive clauses: 338
% Final active clauses: 158
% Final passive clauses: 125
% Main loop iterations started: 201
% 
% >>> Simplifying Inferences
% Duplicate literals: 50
% Fw subsumption resolutions: 76
% Fw demodulations: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 42
% Forward subsumptions: 70
% 
% >>> Generating Inferences
% Binary resolution: 188
% Unit resulting resolution: 3
% Factoring: 1
% Forward superposition: 65
% Backward superposition: 25
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 41
% Split components: 96
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 208
% SAT solver unit clauses: 56
% SAT solver binary clauses: 88
% 
% Memory used [KB]: 5884
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% splitting model update: 0.002 s
% backward superposition index maintenance: 0.001 s
% forward demodulation: 0.001 s
% other: 0.005 s ( own 0.0000 s ) 

