i: [1,45,0,3,1,1,0,0] t28_arytm_1
a: [3,65,0,3,1,45]
s: 65
i: [1,43,0,3,1,1,0,0] t28_arytm_1
a: [3,75,0,3,1,43]
s: 75
i: [1,42,0,3,1,1,0,0] t28_arytm_1
a: [3,80,0,3,1,42]
s: 80
i: [1,41,0,3,1,1,0,0] t28_arytm_1
a: [3,85,0,3,1,41]
s: 85
i: [1,40,0,3,1,1,0,0] t28_arytm_1
a: [3,90,0,3,1,40]
s: 90
i: [1,44,0,5,1,1,0,0] t28_arytm_1
a: [3,70,0,5,1,44]
s: 70
i: [1,47,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,95,0,3,1,47]
s: 95
i: [1,48,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,100,0,3,1,48]
s: 100
i: [1,46,0,15,1,1,0,0] t28_arytm_1
a: [3,60,0,15,1,46]
s: 60
i: [1,52,0,4,1,0,0,1] existence_m1_subset_1
s: 52
i: [1,49,0,8,2,0,0,1] l13_arytm_1
s: 49
i: [1,53,0,11,3,0,0,1] dt_k1_arytm_1
s: 53
i: [1,54,0,11,3,0,0,1] dt_k8_arytm_2
s: 54
d: [2,102,1,12,4,0,44,53,49]
d: [2,103,1,6,2,0,18,102]
a: [3,110,1,3,1,103]
s: 110
i: [1,56,0,12,4,0,0,1] connectedness_r1_arytm_2
s: 56
i: [1,55,0,13,3,0,0,1] commutativity_k8_arytm_2
s: 55
d: [2,115,1,9,3,1,40,56,90]
s: 115
i: [1,50,0,15,5,0,0,1] t4_arytm_1
s: 50
d: [2,163,1,12,4,1,40,50,75]
d: [2,164,1,9,3,2,21,163,85]
d: [2,165,1,6,2,3,21,164,90]
a: [3,169,1,3,1,165]
s: 169
i: [1,51,0,30,6,0,0,1] t27_arytm_1
s: 51
d: [2,175,2,0,0,6,52,90,85,80,169,65,60,51]
e: 175
d: [2,183,1,30,6,1,44,60,51]
d: [2,184,1,15,5,1,19,183]
d: [2,195,1,12,4,2,21,184,65]
d: [2,196,1,9,3,3,21,195,169]
d: [2,197,1,6,2,4,21,196,80]
d: [2,198,1,3,1,5,21,197,85]
d: [2,199,1,0,0,6,21,198,90]
e: 199
r: 169
a: [3,168,1,3,1,165]
s: 168
a: [3,173,1,3,1,165]
s: 173
r: 168
r: 60
r: 70
r: 85
r: 75
d: [2,212,0,5,1,2,23,70,173]
a: [3,224,0,5,1,212]
s: 224
d: [2,227,1,0,0,2,52,90,224,49]
e: 227
d: [2,228,1,6,2,1,44,224,49]
d: [2,229,1,3,1,1,19,228]
d: [2,231,1,0,0,2,21,229,90]
e: 231
d: [2,201,1,9,3,3,21,195,80]
d: [2,202,1,6,2,4,21,201,85]
d: [2,203,1,3,1,5,21,202,90]
d: [2,205,2,12,4,1,40,168,50]
d: [2,206,2,9,3,2,21,205,75]
d: [2,207,2,6,2,3,21,206,90]
d: [2,208,2,3,1,4,21,207,85]
f: 231,212,208,203,40,41,42,43,44,45,46
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t28_arytm_1
% SZS output start Proof for arytm_1__t28_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (r1_arytm_2(X1,X0) => (k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,X0)) = k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(X0,X1))) | k11_arytm_3 = X2 | k1_arytm_1(X0,X1) = k11_arytm_3))))) [input t28_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (r1_arytm_2(X1,X0) => (k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,X0)) = k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(X0,X1))) | k11_arytm_3 = X2 | k1_arytm_1(X0,X1) = k11_arytm_3))))) [negated conjecture 1]
15. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => k11_arytm_3 = k1_arytm_1(X0,X0)) [input l13_arytm_1]
17. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => ~(k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,X0)) != k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(X0,X1))) & k11_arytm_3 != X2 & ~r1_arytm_2(X0,X1))))) [input t27_arytm_1]
18. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((r1_arytm_2(X1,X0) & r1_arytm_2(X0,X1)) => X0 = X1))) [input t4_arytm_1]
19. ? [X0] : (? [X1] : (? [X2] : (((k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,X0)) != k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(X0,X1))) & k11_arytm_3 != X2 & k1_arytm_1(X0,X1) != k11_arytm_3) & r1_arytm_2(X1,X0)) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
20. ? [X0] : (? [X1] : (? [X2] : (k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,X0)) != k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(X0,X1))) & k11_arytm_3 != X2 & k1_arytm_1(X0,X1) != k11_arytm_3 & r1_arytm_2(X1,X0) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 19]
21. ! [X0] : (k11_arytm_3 = k1_arytm_1(X0,X0) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 15]
22. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_arytm_2(X1,X0) | ~r1_arytm_2(X0,X1))) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 18]
23. ! [X0] : (! [X1] : (X0 = X1 | ~r1_arytm_2(X1,X0) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 22]
24. ! [X0] : (! [X1] : (! [X2] : ((k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,X0)) = k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(X0,X1))) | k11_arytm_3 = X2 | r1_arytm_2(X0,X1)) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 17]
25. ! [X0] : (! [X1] : (! [X2] : (k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,X0)) = k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(X0,X1))) | k11_arytm_3 = X2 | r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 24]
34. ? [X0] : (? [X1] : (? [X2] : (k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,X0)) != k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(X0,X1))) & k11_arytm_3 != X2 & k1_arytm_1(X0,X1) != k11_arytm_3 & r1_arytm_2(X1,X0) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (? [X2] : (k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,sK0)) != k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(sK0,X1))) & k11_arytm_3 != X2 & k11_arytm_3 != k1_arytm_1(sK0,X1) & r1_arytm_2(X1,sK0) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
35. ? [X1] : (? [X2] : (k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,sK0)) != k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(sK0,X1))) & k11_arytm_3 != X2 & k11_arytm_3 != k1_arytm_1(sK0,X1) & r1_arytm_2(X1,sK0) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) => (? [X2] : (k2_arytm_1(k8_arytm_2(X2,sK1),k8_arytm_2(X2,sK0)) != k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(sK0,sK1))) & k11_arytm_3 != X2 & k11_arytm_3 != k1_arytm_1(sK0,sK1) & r1_arytm_2(sK1,sK0) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
36. ? [X2] : (k2_arytm_1(k8_arytm_2(X2,sK1),k8_arytm_2(X2,sK0)) != k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(sK0,sK1))) & k11_arytm_3 != X2 & k11_arytm_3 != k1_arytm_1(sK0,sK1) & r1_arytm_2(sK1,sK0) & m1_subset_1(X2,k2_arytm_2)) => (k2_arytm_1(k8_arytm_2(sK2,sK1),k8_arytm_2(sK2,sK0)) != k4_tarski(k11_arytm_3,k8_arytm_2(sK2,k1_arytm_1(sK0,sK1))) & k11_arytm_3 != sK2 & k11_arytm_3 != k1_arytm_1(sK0,sK1) & r1_arytm_2(sK1,sK0) & m1_subset_1(sK2,k2_arytm_2)) [choice axiom]
37. ((k2_arytm_1(k8_arytm_2(sK2,sK1),k8_arytm_2(sK2,sK0)) != k4_tarski(k11_arytm_3,k8_arytm_2(sK2,k1_arytm_1(sK0,sK1))) & k11_arytm_3 != sK2 & k11_arytm_3 != k1_arytm_1(sK0,sK1) & r1_arytm_2(sK1,sK0) & m1_subset_1(sK2,k2_arytm_2)) & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 20,36,35,34]
40. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
41. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
42. m1_subset_1(sK2,k2_arytm_2) [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
43. r1_arytm_2(sK1,sK0) [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
44. k11_arytm_3 != k1_arytm_1(sK0,sK1) [cnf transformation 37] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
45. k11_arytm_3 != sK2 [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
46. k2_arytm_1(k8_arytm_2(sK2,sK1),k8_arytm_2(sK2,sK0)) != k4_tarski(k11_arytm_3,k8_arytm_2(sK2,k1_arytm_1(sK0,sK1))) [cnf transformation 37] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
49. k11_arytm_3 = k1_arytm_1(X0,X0) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 21] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. ~r1_arytm_2(X1,X0) | X0 = X1 | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 23] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k2_arytm_1(k8_arytm_2(X2,X1),k8_arytm_2(X2,X0)) = k4_tarski(k11_arytm_3,k8_arytm_2(X2,k1_arytm_1(X0,X1))) | k11_arytm_3 = X2 | r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 25] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. 1 <=> k2_arytm_1(k8_arytm_2(sK2,sK1),k8_arytm_2(sK2,sK0)) = k4_tarski(k11_arytm_3,k8_arytm_2(sK2,k1_arytm_1(sK0,sK1))) [avatar definition]
60. k2_arytm_1(k8_arytm_2(sK2,sK1),k8_arytm_2(sK2,sK0)) != k4_tarski(k11_arytm_3,k8_arytm_2(sK2,k1_arytm_1(sK0,sK1))) <- (~1) [avatar component clause 58] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
61. ~1 [avatar split clause 46,58]
63. 2 <=> k11_arytm_3 = sK2 [avatar definition]
65. k11_arytm_3 != sK2 <- (~2) [avatar component clause 63] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66. ~2 [avatar split clause 45,63]
68. 3 <=> k11_arytm_3 = k1_arytm_1(sK0,sK1) [avatar definition]
70. k11_arytm_3 != k1_arytm_1(sK0,sK1) <- (~3) [avatar component clause 68] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
71. ~3 [avatar split clause 44,68]
73. 4 <=> r1_arytm_2(sK1,sK0) [avatar definition]
75. r1_arytm_2(sK1,sK0) <- (4) [avatar component clause 73] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
76. 4 [avatar split clause 43,73]
78. 5 <=> m1_subset_1(sK2,k2_arytm_2) [avatar definition]
80. m1_subset_1(sK2,k2_arytm_2) <- (5) [avatar component clause 78] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
81. 5 [avatar split clause 42,78]
83. 6 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
85. m1_subset_1(sK1,k2_arytm_2) <- (6) [avatar component clause 83] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
86. 6 [avatar split clause 41,83]
88. 7 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
90. m1_subset_1(sK0,k2_arytm_2) <- (7) [avatar component clause 88] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
91. 7 [avatar split clause 40,88]
167. 19 <=> r1_arytm_2(sK0,sK1) [avatar definition]
168. r1_arytm_2(sK0,sK1) <- (19) [avatar component clause 167] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
171. 20 <=> sK0 = sK1 [avatar definition]
173. sK0 = sK1 <- (20) [avatar component clause 171] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
183. k2_arytm_1(k8_arytm_2(sK2,sK1),k8_arytm_2(sK2,sK0)) != k2_arytm_1(k8_arytm_2(sK2,sK1),k8_arytm_2(sK2,sK0)) | k11_arytm_3 = sK2 | r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1) [superposition 60,51] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
184. k11_arytm_3 = sK2 | r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1) [trivial inequality removal 183] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
195. r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1, ~2) [subsumption resolution 184,65] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
201. r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1, ~2, 5) [subsumption resolution 195,80] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
202. r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1, ~2, 5, 6) [subsumption resolution 201,85] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
203. r1_arytm_2(sK0,sK1) <- (~1, ~2, 5, 6, 7) [subsumption resolution 202,90] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
204. 19 | 1 | 2 | ~5 | ~6 | ~7 [avatar split clause 203,88,83,78,63,58,167]
205. sK0 = sK1 | ~r1_arytm_2(sK1,sK0) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (19) [resolution 168,50] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
206. sK0 = sK1 | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (4, 19) [subsumption resolution 205,75] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
207. sK0 = sK1 | ~m1_subset_1(sK1,k2_arytm_2) <- (4, 7, 19) [subsumption resolution 206,90] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
208. sK0 = sK1 <- (4, 6, 7, 19) [subsumption resolution 207,85] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
209. 20 | ~4 | ~6 | ~7 | ~19 [avatar split clause 208,167,88,83,73,171]
212. k11_arytm_3 != k1_arytm_1(sK0,sK0) <- (~3, 20) [backward demodulation 70,173] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
222. 22 <=> k11_arytm_3 = k1_arytm_1(sK0,sK0) [avatar definition]
224. k11_arytm_3 != k1_arytm_1(sK0,sK0) <- (~22) [avatar component clause 222] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
225. ~22 | 3 | ~20 [avatar split clause 212,171,68,222]
228. k11_arytm_3 != k11_arytm_3 | ~m1_subset_1(sK0,k2_arytm_2) <- (~22) [superposition 224,49] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
229. ~m1_subset_1(sK0,k2_arytm_2) <- (~22) [trivial inequality removal 228] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
231. $false <- (7, ~22) [subsumption resolution 229,90] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
232. ~7 | 22 [avatar contradiction clause 231]
233. $false [avatar sat refutation 61,66,71,76,81,86,91,204,209,225,232] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:194,thDist:-194}
% SZS output end Proof for arytm_1__t28_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 108
% Active clauses: 23
% Passive clauses: 43
% Final active clauses: 17
% Final passive clauses: 20
% Main loop iterations started: 23
% 
% >>> Simplifying Inferences
% Duplicate literals: 19
% Trivial inequalities: 2
% Fw subsumption resolutions: 16
% Bw demodulations: 5
% 
% >>> Deletion Inferences
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 16
% Unit resulting resolution: 2
% Forward superposition: 9
% Backward superposition: 3
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 27
% SAT solver unit clauses: 9
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

