i: [1,39,0,2,1,1,0,0] t11_sysrel
a: [3,68,0,2,1,39]
s: 68
i: [1,38,0,2,1,1,0,0] t11_sysrel
a: [3,73,0,2,1,38]
s: 73
i: [1,40,0,3,1,1,0,0] t11_sysrel
a: [3,63,0,3,1,40]
s: 63
i: [1,41,0,5,1,1,0,0] t11_sysrel
a: [3,58,0,5,1,41]
s: 58
i: [1,51,0,3,1,0,0,1] reflexivity_r1_tarski
s: 51
i: [1,50,0,4,1,0,0,255] existence_m1_subset_1
s: 50
i: [1,42,0,5,2,0,0,1] dt_k2_relat_1
s: 42
i: [1,46,0,5,1,0,0,1] t7_xboole_1
s: 46
i: [1,52,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 52
i: [1,43,0,7,2,0,0,1] involutiveness_k2_relat_1
s: 43
i: [1,48,0,7,2,0,0,255] t3_subset
s: 48
d: [2,77,1,5,1,0,40,48,50]
s: 77
i: [1,49,0,7,2,0,0,255] t3_subset
s: 49
i: [1,53,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 53
i: [1,45,0,8,3,0,0,255] cc2_relat_1
s: 45
i: [1,47,0,8,2,0,0,1] t12_xboole_1
s: 47
d: [2,79,1,5,1,0,44,46,53]
s: 79
i: [1,54,0,8,3,0,0,1] fc3_relat_1
s: 54
d: [2,80,1,5,1,0,44,46,53]
s: 80
i: [1,44,0,14,3,0,0,1] t23_relat_1
s: 44
d: [2,81,1,6,2,0,40,45,50]
s: 81
d: [2,82,1,7,3,0,40,45,49]
s: 82
d: [2,83,1,8,2,0,44,47,53]
s: 83
d: [2,118,2,6,2,0,40,82,46]
s: 118
d: [2,84,1,8,2,0,44,47,53]
s: 84
d: [2,85,1,8,2,0,44,53,47]
s: 85
d: [2,119,2,6,2,0,40,82,79]
s: 119
d: [2,86,1,8,2,0,44,53,47]
s: 86
d: [2,136,3,6,2,0,44,118,53]
s: 136
d: [2,103,1,15,5,0,44,54,44]
d: [2,114,1,12,4,0,21,103,42]
d: [2,115,1,9,3,0,21,114,42]
s: 115
d: [2,108,1,11,3,0,44,46,44]
s: 108
d: [2,137,3,6,2,0,44,118,53]
s: 137
d: [2,101,1,14,3,0,44,44,53]
s: 101
d: [2,121,2,9,3,0,44,83,47]
s: 121
d: [2,102,1,14,3,0,44,44,53]
s: 102
d: [2,106,1,14,3,0,44,53,44]
s: 106
d: [2,203,3,6,2,1,40,121,63]
a: [3,208,3,3,1,203]
s: 208
d: [2,107,1,14,3,0,44,53,44]
s: 107
d: [2,126,2,9,3,0,44,47,83]
s: 126
d: [2,97,1,17,4,0,44,44,43]
d: [2,112,1,14,3,0,21,97,42]
s: 112
d: [2,98,1,17,4,0,44,44,43]
d: [2,113,1,14,3,0,21,98,42]
s: 113
d: [2,224,2,14,5,0,44,115,112]
d: [2,232,2,12,4,0,18,224]
d: [2,241,2,9,3,0,21,232,42]
s: 241
d: [2,100,1,16,4,0,44,44,47]
s: 100
d: [2,261,2,14,5,0,44,115,113]
d: [2,271,2,12,4,0,18,261]
d: [2,283,2,9,3,0,21,271,42]
s: 283
d: [2,105,1,16,4,0,44,47,44]
s: 105
d: [2,151,2,16,5,0,44,115,44]
d: [2,153,2,13,4,0,21,151,42]
d: [2,154,2,10,3,0,21,153,42]
s: 154
d: [2,378,3,8,3,0,44,154,52]
d: [2,392,3,6,2,0,18,378]
s: 392
d: [2,159,2,12,4,0,44,108,47]
d: [2,165,2,10,3,0,21,159,82]
s: 165
d: [2,409,3,0,0,3,52,68,63,58,165]
e: 409
d: [2,410,3,5,2,1,40,165,58]
d: [2,422,3,3,1,2,21,410,68]
d: [2,423,3,0,0,3,21,422,63]
e: 423
f: 423,39,40,41
% Refutation found. Thanks to Tanya!
% SZS status Theorem for sysrel__t11_sysrel
% SZS output start Proof for sysrel__t11_sysrel
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,X1) => r1_tarski(k2_relat_1(X0),k2_relat_1(X1))))) [input t11_sysrel]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,X1) => r1_tarski(k2_relat_1(X0),k2_relat_1(X1))))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
14. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k2_xboole_0(X0,X1) = X1) [input t12_xboole_1]
15. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => k2_relat_1(k2_xboole_0(X0,X1)) = k2_xboole_0(k2_relat_1(X0),k2_relat_1(X1)))) [input t23_relat_1]
16. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
17. ! [X0] : ! [X1] : r1_tarski(X0,k2_xboole_0(X0,X1)) [input t7_xboole_1]
18. ! [X0,X1] : r1_tarski(X0,k2_xboole_0(X0,X1)) [flattening 17]
19. ! [X0,X1] : (r1_tarski(X0,X1) => k2_xboole_0(X0,X1) = X1) [flattening 14]
20. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 16]
23. ? [X0] : (? [X1] : ((~r1_tarski(k2_relat_1(X0),k2_relat_1(X1)) & r1_tarski(X0,X1)) & v1_relat_1(X1)) & v1_relat_1(X0)) [ennf transformation 2]
24. ? [X0] : (? [X1] : (~r1_tarski(k2_relat_1(X0),k2_relat_1(X1)) & r1_tarski(X0,X1) & v1_relat_1(X1)) & v1_relat_1(X0)) [flattening 23]
27. ! [X0] : (! [X1] : (k2_relat_1(k2_xboole_0(X0,X1)) = k2_xboole_0(k2_relat_1(X0),k2_relat_1(X1)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 15]
28. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
29. ! [X0,X1] : (k2_xboole_0(X0,X1) = X1 | ~r1_tarski(X0,X1)) [ennf transformation 19]
32. ? [X0] : (? [X1] : (~r1_tarski(k2_relat_1(X0),k2_relat_1(X1)) & r1_tarski(X0,X1) & v1_relat_1(X1)) & v1_relat_1(X0)) => (? [X1] : (~r1_tarski(k2_relat_1(sK0),k2_relat_1(X1)) & r1_tarski(sK0,X1) & v1_relat_1(X1)) & v1_relat_1(sK0)) [choice axiom]
33. ? [X1] : (~r1_tarski(k2_relat_1(sK0),k2_relat_1(X1)) & r1_tarski(sK0,X1) & v1_relat_1(X1)) => (~r1_tarski(k2_relat_1(sK0),k2_relat_1(sK1)) & r1_tarski(sK0,sK1) & v1_relat_1(sK1)) [choice axiom]
34. (~r1_tarski(k2_relat_1(sK0),k2_relat_1(sK1)) & r1_tarski(sK0,sK1) & v1_relat_1(sK1)) & v1_relat_1(sK0) [skolemisation 24,33,32]
35. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 20]
39. v1_relat_1(sK1) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. r1_tarski(sK0,sK1) [cnf transformation 34] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. ~r1_tarski(k2_relat_1(sK0),k2_relat_1(sK1)) [cnf transformation 34] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
44. k2_relat_1(k2_xboole_0(X0,X1)) = k2_xboole_0(k2_relat_1(X0),k2_relat_1(X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
46. r1_tarski(X0,k2_xboole_0(X0,X1)) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. k2_xboole_0(X0,X1) = X1 | ~r1_tarski(X0,X1) [cnf transformation 29] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 35] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
56. 1 <=> r1_tarski(k2_relat_1(sK0),k2_relat_1(sK1)) [avatar definition]
58. ~r1_tarski(k2_relat_1(sK0),k2_relat_1(sK1)) <- (~1) [avatar component clause 56] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
59. ~1 [avatar split clause 41,56]
61. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
63. r1_tarski(sK0,sK1) <- (2) [avatar component clause 61] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
64. 2 [avatar split clause 40,61]
66. 3 <=> v1_relat_1(sK1) [avatar definition]
68. v1_relat_1(sK1) <- (3) [avatar component clause 66] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
69. 3 [avatar split clause 39,66]
82. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 45,49] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
108. r1_tarski(k2_relat_1(X10),k2_relat_1(k2_xboole_0(X10,X11))) | ~v1_relat_1(X11) | ~v1_relat_1(X10) [superposition 46,44] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
159. r1_tarski(k2_relat_1(X3),k2_relat_1(X4)) | ~v1_relat_1(X4) | ~v1_relat_1(X3) | ~r1_tarski(X3,X4) [superposition 108,47] {a:2,w:12,wCS:4800,sine:1,thAx:0,allAx:3,thDist:-3}
165. r1_tarski(k2_relat_1(X3),k2_relat_1(X4)) | ~v1_relat_1(X4) | ~r1_tarski(X3,X4) [subsumption resolution 159,82] {a:2,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
410. ~v1_relat_1(sK1) | ~r1_tarski(sK0,sK1) <- (~1) [resolution 165,58] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
422. ~r1_tarski(sK0,sK1) <- (~1, 3) [subsumption resolution 410,68] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
423. $false <- (~1, 2, 3) [subsumption resolution 422,63] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
424. 1 | ~2 | ~3 [avatar contradiction clause 423]
431. $false [avatar sat refutation 59,64,69,424] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:43,thDist:-43}
% SZS output end Proof for sysrel__t11_sysrel
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 367
% Active clauses: 37
% Passive clauses: 136
% Final active clauses: 37
% Final passive clauses: 88
% Main loop iterations started: 48
% 
% >>> Simplifying Inferences
% Duplicate literals: 37
% Fw subsumption resolutions: 93
% Fw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 24
% Equational tautologies: 1
% Forward subsumptions: 70
% Fw demodulations to eq. taut.: 4
% 
% >>> Generating Inferences
% Binary resolution: 25
% Unit resulting resolution: 1
% Forward superposition: 123
% Backward superposition: 68
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 4
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

