i: [1,170,0,2,1,1,0,0] t5_int_2
a: [3,265,0,2,1,170]
s: 265
i: [1,169,0,2,1,1,0,0] t5_int_2
a: [3,270,0,2,1,169]
s: 270
i: [1,173,0,11,3,1,0,0] t5_int_2
a: [3,250,0,3,1,173]
s: 250
a: [3,257,0,5,1,173]
s: 257
i: [1,174,0,2,1,0,0,2] fc1_xboole_0
a: [3,275,0,2,1,174]
s: 275
i: [1,181,0,2,1,0,0,3] fc6_ordinal1
a: [3,305,0,2,1,181]
s: 305
i: [1,180,0,2,1,0,0,3] fc6_ordinal1
a: [3,310,0,2,1,180]
s: 310
i: [1,243,0,5,2,0,0,255] t3_int_2
a: [3,315,0,2,1,243]
s: 315
i: [1,222,0,7,3,0,0,255] reflexivity_r1_int_1
a: [3,323,0,2,1,222]
s: 323
r: 315
d: [2,420,1,0,0,2,40,323,270]
e: 420
d: [2,419,1,0,0,2,40,323,265]
e: 419
d: [2,418,1,0,0,2,52,265,323]
e: 418
r: 323
s: 315
i: [1,225,0,2,1,0,0,2] rc2_xboole_0
a: [3,331,0,2,1,225]
s: 331
i: [1,226,0,2,1,0,0,1] rc2_int_1
a: [3,336,0,2,1,226]
s: 336
i: [1,227,0,2,1,0,0,2] rc4_ordinal1
a: [3,341,0,2,1,227]
s: 341
i: [1,228,0,2,1,0,0,3] rc1_ordinal1
a: [3,346,0,2,1,228]
s: 346
i: [1,229,0,2,1,0,0,2] rc1_xboole_0
a: [3,351,0,2,1,229]
s: 351
i: [1,233,0,2,1,0,0,255] rc2_ordinal1
a: [3,356,0,2,1,233]
s: 356
i: [1,232,0,2,1,0,0,255] rc2_ordinal1
a: [3,361,0,2,1,232]
s: 361
i: [1,231,0,2,1,0,0,255] rc2_ordinal1
a: [3,366,0,2,1,231]
s: 366
i: [1,230,0,2,1,0,0,255] rc2_ordinal1
a: [3,371,0,2,1,230]
s: 371
i: [1,235,0,2,1,0,0,2] rc5_ordinal1
a: [3,376,0,2,1,235]
s: 376
i: [1,234,0,2,1,0,0,2] rc5_ordinal1
a: [3,381,0,2,1,234]
s: 381
i: [1,236,0,2,1,0,0,255] rc3_xreal_0
a: [3,386,0,2,1,236]
s: 386
i: [1,237,0,2,1,0,0,255] rc4_xreal_0
a: [3,391,0,2,1,237]
s: 391
i: [1,239,0,2,1,0,0,255] rc1_int_1
a: [3,396,0,2,1,239]
s: 396
i: [1,242,0,2,1,0,0,2] rc2_nat_1
a: [3,406,0,2,1,242]
s: 406
i: [1,241,0,2,1,0,0,2] rc2_nat_1
a: [3,411,0,2,1,241]
s: 411
i: [1,216,0,3,1,0,0,255] reflexivity_r1_tarski
s: 216
i: [1,175,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,280,0,3,1,175]
s: 280
r: 310
r: 305
d: [2,424,0,2,1,2,23,310,280]
a: [3,429,0,2,1,424]
s: 429
d: [2,425,0,2,1,2,23,305,280]
a: [3,434,0,2,1,425]
s: 434
i: [1,176,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,285,0,3,1,176]
s: 285
r: 275
d: [2,436,0,2,1,2,23,275,285]
a: [3,440,0,2,1,436]
s: 440
i: [1,177,0,3,1,0,0,2] t1_numerals
a: [3,290,0,3,1,177]
s: 290
i: [1,238,0,3,1,0,0,255] rc1_int_1
a: [3,401,0,3,1,238]
s: 401
d: [2,442,0,3,1,2,22,290,280]
d: [2,443,0,3,1,3,22,442,285]
a: [3,447,0,3,1,443]
s: 447
i: [1,192,0,4,2,0,0,1] cc2_int_1
s: 192
d: [2,449,1,2,1,1,40,192,341]
a: [3,454,1,2,1,449]
s: 454
i: [1,193,0,4,2,0,0,3] cc6_ordinal1
s: 193
d: [2,450,1,2,1,1,40,192,376]
a: [3,459,1,2,1,450]
s: 459
i: [1,195,0,4,2,0,0,255] cc3_nat_1
s: 195
i: [1,196,0,4,2,0,0,3] cc1_nat_1
s: 196
d: [2,461,1,2,1,1,40,193,341]
a: [3,466,1,2,1,461]
s: 466
i: [1,198,0,4,2,0,0,255] cc1_ordinal1
s: 198
d: [2,462,1,2,1,1,40,193,376]
a: [3,471,1,2,1,462]
s: 471
i: [1,199,0,4,2,0,0,255] cc1_ordinal1
s: 199
i: [1,201,0,4,2,0,0,255] cc9_ordinal1
s: 201
d: [2,473,1,2,1,1,40,195,386]
a: [3,477,1,2,1,473]
s: 477
i: [1,202,0,4,2,0,0,2] cc7_ordinal1
s: 202
d: [2,479,1,2,1,1,40,198,434]
a: [3,487,1,2,1,479]
s: 487
i: [1,203,0,4,2,0,0,3] cc3_ordinal1
s: 203
i: [1,215,0,4,1,0,0,3] existence_m1_subset_1
s: 215
d: [2,481,1,2,1,1,40,198,346]
a: [3,497,1,2,1,481]
s: 497
i: [1,178,0,4,1,0,0,2] dt_k5_numbers
a: [3,295,0,4,1,178]
s: 295
d: [2,483,1,2,1,1,40,198,406]
a: [3,502,1,2,1,483]
s: 502
i: [1,179,0,4,1,0,0,1] dt_k6_numbers
a: [3,300,0,4,1,179]
s: 300
i: [1,240,0,4,1,0,0,2] rc2_nat_1
a: [3,416,0,4,1,240]
s: 416
d: [2,510,1,2,1,1,40,199,434]
a: [3,519,1,2,1,510]
s: 519
i: [1,204,0,5,2,0,0,2] t6_boole
s: 204
d: [2,512,1,2,1,1,40,199,346]
a: [3,529,1,2,1,512]
s: 529
i: [1,205,0,5,2,0,0,255] cc2_nat_1
s: 205
i: [1,206,0,5,2,0,0,3] cc8_ordinal1
s: 206
d: [2,515,1,2,1,1,40,199,406]
a: [3,539,1,2,1,515]
s: 539
i: [1,212,0,5,2,0,0,255] t7_boole
s: 212
d: [2,541,1,2,1,1,40,201,391]
a: [3,547,1,2,1,541]
s: 547
a: [3,326,0,5,2,222]
s: 326
d: [2,595,0,5,2,1,22,204,285]
s: 595
d: [2,542,1,2,1,1,40,201,351]
a: [3,552,1,2,1,542]
s: 552
d: [2,608,0,5,2,1,22,206,280]
s: 608
d: [2,543,1,2,1,1,40,201,440]
a: [3,557,1,2,1,543]
s: 557
i: [1,207,0,6,3,0,0,255] cc2_ordinal1
s: 207
i: [1,208,0,6,2,0,0,255] t1_subset
s: 208
d: [2,559,1,2,1,1,40,202,391]
a: [3,565,1,2,1,559]
s: 565
i: [1,217,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 217
d: [2,560,1,2,1,1,40,202,351]
a: [3,570,1,2,1,560]
s: 570
i: [1,200,0,7,3,0,0,3] cc5_ordinal1
s: 200
i: [1,210,0,7,2,0,0,255] t3_subset
s: 210
d: [2,561,1,2,1,1,40,202,440]
a: [3,575,1,2,1,561]
s: 575
d: [2,660,2,0,0,2,52,315,575,192]
e: 660
d: [2,662,2,2,1,1,40,575,192]
d: [2,665,2,0,0,2,21,662,315]
e: 665
r: 315
a: [3,314,0,2,1,243]
s: 314
d: [2,577,1,2,1,1,40,203,391]
a: [3,583,1,2,1,577]
s: 583
a: [3,319,0,3,1,243]
s: 319
i: [1,211,0,7,3,0,0,1] t8_boole
s: 211
d: [2,578,1,2,1,1,40,203,351]
a: [3,588,1,2,1,578]
s: 588
i: [1,182,0,8,3,0,0,255] cc10_ordinal1
s: 182
d: [2,579,1,2,1,1,40,203,440]
a: [3,593,1,2,1,579]
s: 593
i: [1,183,0,8,3,0,0,255] t3_int_2
s: 183
i: [1,209,0,8,3,0,0,255] t2_subset
s: 209
d: [2,596,1,2,1,1,40,205,447]
a: [3,601,1,2,1,596]
s: 601
i: [1,223,0,8,3,0,0,1] dt_k3_int_2
s: 223
d: [2,646,1,4,2,1,40,200,401]
a: [3,653,1,2,1,646]
s: 653
i: [1,214,0,9,3,0,0,255] t5_subset
s: 214
i: [1,213,0,10,3,0,0,255] t4_subset
s: 213
d: [2,696,1,4,2,1,40,182,295]
a: [3,702,1,2,1,696]
s: 702
i: [1,224,0,11,3,0,0,1] commutativity_k3_int_2
s: 224
d: [2,734,1,5,2,1,40,209,401]
a: [3,757,1,2,1,734]
s: 757
i: [1,245,0,13,4,0,0,255] d2_int_2
s: 245
d: [2,789,0,9,3,0,21,245,223]
s: 789
d: [2,790,1,7,3,1,44,789,257]
d: [2,795,1,5,2,2,21,790,265]
d: [2,796,1,3,1,3,21,795,270]
a: [3,800,1,3,1,796]
s: 800
a: [3,253,0,3,1,173]
s: 253
r: 800
r: 257
r: 265
d: [2,807,0,5,1,2,23,257,253]
a: [3,812,0,5,1,807]
s: 812
i: [1,246,0,13,4,0,0,255] d2_int_2
s: 246
d: [2,837,0,9,3,0,21,246,223]
s: 837
d: [2,840,1,7,3,1,44,837,812]
d: [2,843,1,5,2,2,21,840,314]
d: [2,844,1,3,1,3,21,843,270]
a: [3,848,1,3,1,844]
s: 848
d: [2,850,2,0,0,3,52,270,250,848,183]
e: 850
d: [2,851,2,5,2,1,40,848,183]
d: [2,854,2,2,1,2,21,851,250]
d: [2,855,2,0,0,3,21,854,270]
e: 855
r: 250
r: 812
r: 848
a: [3,249,0,3,1,173]
s: 249
r: 270
a: [3,811,0,5,1,807]
s: 811
d: [2,861,0,5,1,2,22,811,249]
a: [3,865,0,5,1,861]
s: 865
i: [1,221,0,14,4,0,0,2] existence_m2_subset_1
s: 221
d: [2,815,1,9,3,1,44,812,224]
d: [2,821,1,7,2,2,21,815,270]
d: [2,822,1,5,1,3,21,821,314]
a: [3,825,1,5,1,822]
s: 825
i: [1,218,0,15,5,0,0,2] dt_m2_subset_1
s: 218
d: [2,803,2,5,2,2,44,800,595]
s: 803
i: [1,219,0,15,5,0,0,2] redefinition_m2_subset_1
s: 219
i: [1,220,0,15,5,0,0,2] redefinition_m2_subset_1
s: 220
d: [2,880,2,5,2,3,22,803,253]
s: 880
i: [1,244,0,21,7,0,0,255] d2_int_2
s: 244
d: [2,867,1,7,2,2,44,865,595]
s: 867
d: [2,894,0,17,6,0,21,244,223]
s: 894
i: [1,188,0,22,7,0,0,255] d2_int_2
s: 188
d: [2,893,3,7,3,3,44,880,595]
s: 893
i: [1,189,0,23,7,0,0,255] d2_int_2
s: 189
d: [2,480,2,2,1,1,40,198,466]
a: [3,492,2,2,1,480]
s: 492
i: [1,190,0,23,7,0,0,255] d2_int_2
s: 190
i: [1,191,0,23,7,0,0,255] d2_int_2
s: 191
d: [2,504,2,2,1,1,40,471,198]
a: [3,508,2,2,1,504]
s: 508
d: [2,597,1,3,1,0,40,205,215]
a: [3,606,1,3,1,597]
s: 606
d: [2,511,2,2,1,1,40,199,466]
a: [3,524,2,2,1,511]
s: 524
d: [2,615,1,3,1,1,40,608,215]
a: [3,620,1,3,1,615]
s: 620
d: [2,731,1,5,2,1,40,209,447]
d: [2,736,1,3,1,2,21,731,429]
a: [3,740,1,3,1,736]
s: 740
d: [2,514,2,2,1,1,40,199,471]
a: [3,534,2,2,1,514]
s: 534
d: [2,732,1,7,2,1,40,209,295]
a: [3,749,1,3,1,732]
s: 749
r: 749
d: [2,628,2,2,1,1,40,565,192]
a: [3,633,2,2,1,628]
s: 633
a: [3,748,1,3,1,732]
s: 748
a: [3,754,1,3,1,734]
s: 754
d: [2,636,2,2,1,1,40,570,192]
a: [3,641,2,2,1,636]
s: 641
d: [2,869,1,11,4,1,40,218,300]
d: [2,872,1,7,3,2,21,869,295]
d: [2,873,1,5,2,3,21,872,429]
d: [2,874,1,3,1,4,21,873,757]
a: [3,878,1,3,1,874]
s: 878
d: [2,670,2,2,1,1,40,583,199]
a: [3,675,2,2,1,670]
s: 675
d: [2,669,1,4,2,2,44,314,595]
s: 669
a: [3,745,1,4,1,732]
s: 745
d: [2,671,2,2,1,1,40,583,198]
a: [3,680,2,2,1,671]
s: 680
d: [2,735,1,7,2,1,40,209,416]
a: [3,763,1,4,1,735]
s: 763
d: [2,683,2,2,1,1,40,588,199]
a: [3,688,2,2,1,683]
s: 688
d: [2,611,1,5,2,2,44,447,595]
s: 611
d: [2,645,1,5,2,0,40,200,215]
s: 645
d: [2,684,2,2,1,1,40,588,198]
a: [3,693,2,2,1,684]
s: 693
d: [2,682,1,5,2,2,44,319,595]
s: 682
d: [2,713,2,2,1,1,40,593,199]
a: [3,719,2,2,1,713]
s: 719
d: [2,610,1,6,2,2,44,300,595]
s: 610
d: [2,697,1,6,2,0,40,182,215]
s: 697
d: [2,714,2,2,1,1,40,593,198]
a: [3,724,2,2,1,714]
s: 724
d: [2,733,1,6,2,0,40,209,215]
s: 733
d: [2,915,2,3,1,1,40,620,193]
a: [3,920,2,3,1,915]
s: 920
d: [2,782,1,6,2,1,40,213,295]
s: 782
d: [2,784,1,6,2,1,40,213,416]
s: 784
d: [2,916,2,3,1,1,40,620,192]
a: [3,925,2,3,1,916]
s: 925
d: [2,609,1,7,3,1,44,595,595]
s: 609
d: [2,927,2,3,1,1,40,740,217]
a: [3,934,2,3,1,927]
s: 934
d: [2,695,1,7,3,0,40,182,210]
s: 695
d: [2,771,1,7,2,0,40,214,215]
s: 771
d: [2,955,2,3,1,1,40,754,217]
a: [3,961,2,3,1,955]
s: 961
d: [2,885,1,11,4,1,40,220,295]
d: [2,888,1,9,3,2,21,885,429]
d: [2,889,1,7,2,3,21,888,757]
s: 889
d: [2,963,2,5,2,1,40,878,209]
d: [2,966,2,3,1,2,21,963,757]
a: [3,970,2,3,1,966]
s: 970
d: [2,887,1,11,4,1,40,220,416]
d: [2,890,1,9,3,2,21,887,411]
d: [2,891,1,7,2,3,21,890,757]
s: 891
d: [2,659,1,8,3,0,40,210,200]
s: 659
d: [2,1012,3,3,1,1,40,920,199]
a: [3,1017,3,3,1,1012]
s: 1017
d: [2,766,1,8,3,0,40,223,193]
s: 766
d: [2,1013,3,3,1,1,40,920,198]
a: [3,1022,3,3,1,1013]
s: 1022
d: [2,767,1,8,3,0,40,223,192]
s: 767
d: [2,769,1,8,3,0,40,214,210]
s: 769
d: [2,1047,3,3,1,1,40,970,217]
a: [3,1054,3,3,1,1047]
s: 1054
d: [2,783,1,8,2,0,40,213,215]
s: 783
d: [2,765,2,4,2,2,44,601,595]
s: 765
d: [2,781,1,9,3,0,40,213,210]
s: 781
d: [2,792,1,13,5,0,44,789,224]
d: [2,793,1,9,3,0,18,792]
s: 793
d: [2,1077,3,2,1,3,40,765,351]
a: [3,1082,3,2,1,1077]
s: 1082
d: [2,791,1,13,5,0,44,789,224]
d: [2,794,1,9,3,0,18,791]
s: 794
d: [2,1078,3,2,1,3,40,765,391]
a: [3,1087,3,2,1,1078]
s: 1087
d: [2,728,1,10,4,1,44,183,595]
s: 728
d: [2,730,1,10,3,0,40,209,210]
s: 730
d: [2,976,2,4,1,1,40,745,217]
a: [3,982,2,4,1,976]
s: 982
d: [2,870,1,21,7,0,40,218,221]
d: [2,871,1,13,4,0,18,870]
s: 871
d: [2,984,2,4,1,1,40,763,217]
a: [3,990,2,4,1,984]
s: 990
d: [2,882,1,21,7,0,40,219,221]
d: [2,883,1,13,4,0,18,882]
s: 883
d: [2,884,1,14,5,0,40,220,210]
s: 884
d: [2,994,2,4,2,2,40,611,205]
s: 994
d: [2,902,1,34,12,0,40,191,189]
d: [2,908,1,17,6,0,18,902]
d: [2,909,1,15,5,0,21,908,192]
s: 909
d: [2,1140,2,0,0,4,52,314,575,319,319,865,909]
e: 1140
d: [2,1152,2,13,5,1,44,865,909]
d: [2,1153,2,10,4,1,18,1152]
d: [2,1154,2,7,3,1,19,1153]
d: [2,1175,2,4,2,2,21,1154,319]
d: [2,1176,2,2,1,3,21,1175,575]
d: [2,1177,2,0,0,4,21,1176,314]
e: 1177
i: [1,172,0,8,2,1,0,0] t5_int_2
i: [1,171,0,8,2,1,0,0] t5_int_2
d: [2,614,1,2,1,2,40,608,447]
d: [2,802,2,5,2,1,40,800,183]
d: [2,804,2,3,1,2,21,802,265]
a: [3,258,0,5,1,173]
d: [2,858,0,5,1,2,22,258,253]
f: 1177,861,858,855,844,807,804,796,662,614,443,243,177,176,175,169,170,171,172,173
% Refutation found. Thanks to Tanya!
% SZS status Theorem for int_2__t5_int_2
% SZS output start Proof for int_2__t5_int_2
1. ! [X0] : (v1_int_1(X0) => ! [X1] : (v1_int_1(X1) => ((k6_numbers = X1 & k6_numbers = X0) <=> k6_numbers = k3_int_2(X0,X1)))) [input t5_int_2]
2. ~! [X0] : (v1_int_1(X0) => ! [X1] : (v1_int_1(X1) => ((k6_numbers = X1 & k6_numbers = X0) <=> k6_numbers = k3_int_2(X0,X1)))) [negated conjecture 1]
8. ! [X0] : (v7_ordinal1(X0) => v1_int_1(X0)) [input cc2_int_1]
21. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
24. ! [X0] : (v1_int_1(X0) => ! [X1] : (v1_int_1(X1) => ! [X2] : (v7_ordinal1(X2) => (k3_int_2(X0,X1) = X2 <=> (! [X3] : (v1_int_1(X3) => ((r1_int_1(X3,X1) & r1_int_1(X3,X0)) => r1_int_1(X3,X2))) & r1_int_1(X2,X1) & r1_int_1(X2,X0)))))) [input d2_int_2]
28. ! [X0,X1] : ((v1_int_1(X1) & v1_int_1(X0)) => v7_ordinal1(k3_int_2(X0,X1))) [input dt_k3_int_2]
50. k5_numbers = k4_ordinal1 [input redefinition_k5_numbers]
51. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
55. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
58. ! [X0] : (v1_int_1(X0) => (r1_int_1(k6_numbers,X0) <=> k6_numbers = X0)) [input t3_int_2]
90. ? [X0] : (? [X1] : (((k6_numbers = X1 & k6_numbers = X0) <~> k6_numbers = k3_int_2(X0,X1)) & v1_int_1(X1)) & v1_int_1(X0)) [ennf transformation 2]
92. ! [X0] : ((r1_int_1(k6_numbers,X0) <=> k6_numbers = X0) | ~v1_int_1(X0)) [ennf transformation 58]
93. ! [X0] : (! [X1] : (! [X2] : ((k3_int_2(X0,X1) = X2 <=> (! [X3] : ((r1_int_1(X3,X2) | (~r1_int_1(X3,X1) | ~r1_int_1(X3,X0))) | ~v1_int_1(X3)) & r1_int_1(X2,X1) & r1_int_1(X2,X0))) | ~v7_ordinal1(X2)) | ~v1_int_1(X1)) | ~v1_int_1(X0)) [ennf transformation 24]
94. ! [X0] : (! [X1] : (! [X2] : ((k3_int_2(X0,X1) = X2 <=> (! [X3] : (r1_int_1(X3,X2) | ~r1_int_1(X3,X1) | ~r1_int_1(X3,X0) | ~v1_int_1(X3)) & r1_int_1(X2,X1) & r1_int_1(X2,X0))) | ~v7_ordinal1(X2)) | ~v1_int_1(X1)) | ~v1_int_1(X0)) [flattening 93]
95. ! [X0] : (v1_int_1(X0) | ~v7_ordinal1(X0)) [ennf transformation 8]
106. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 21]
127. ! [X0,X1] : (v7_ordinal1(k3_int_2(X0,X1)) | (~v1_int_1(X1) | ~v1_int_1(X0))) [ennf transformation 28]
128. ! [X0,X1] : (v7_ordinal1(k3_int_2(X0,X1)) | ~v1_int_1(X1) | ~v1_int_1(X0)) [flattening 127]
131. ? [X0] : (? [X1] : (((k6_numbers != k3_int_2(X0,X1) | (k6_numbers != X1 | k6_numbers != X0)) & (k6_numbers = k3_int_2(X0,X1) | (k6_numbers = X1 & k6_numbers = X0))) & v1_int_1(X1)) & v1_int_1(X0)) [nnf transformation 90]
132. ? [X0] : (? [X1] : ((k6_numbers != k3_int_2(X0,X1) | k6_numbers != X1 | k6_numbers != X0) & (k6_numbers = k3_int_2(X0,X1) | (k6_numbers = X1 & k6_numbers = X0)) & v1_int_1(X1)) & v1_int_1(X0)) [flattening 131]
133. ? [X0] : (? [X1] : ((k6_numbers != k3_int_2(X0,X1) | k6_numbers != X1 | k6_numbers != X0) & (k6_numbers = k3_int_2(X0,X1) | (k6_numbers = X1 & k6_numbers = X0)) & v1_int_1(X1)) & v1_int_1(X0)) => (? [X1] : ((k6_numbers != k3_int_2(sK0,X1) | k6_numbers != X1 | k6_numbers != sK0) & (k6_numbers = k3_int_2(sK0,X1) | (k6_numbers = X1 & k6_numbers = sK0)) & v1_int_1(X1)) & v1_int_1(sK0)) [choice axiom]
134. ? [X1] : ((k6_numbers != k3_int_2(sK0,X1) | k6_numbers != X1 | k6_numbers != sK0) & (k6_numbers = k3_int_2(sK0,X1) | (k6_numbers = X1 & k6_numbers = sK0)) & v1_int_1(X1)) => ((k6_numbers != k3_int_2(sK0,sK1) | k6_numbers != sK1 | k6_numbers != sK0) & (k6_numbers = k3_int_2(sK0,sK1) | (k6_numbers = sK1 & k6_numbers = sK0)) & v1_int_1(sK1)) [choice axiom]
135. ((k6_numbers != k3_int_2(sK0,sK1) | k6_numbers != sK1 | k6_numbers != sK0) & (k6_numbers = k3_int_2(sK0,sK1) | (k6_numbers = sK1 & k6_numbers = sK0)) & v1_int_1(sK1)) & v1_int_1(sK0) [skolemisation 132,134,133]
136. ! [X0] : (((r1_int_1(k6_numbers,X0) | k6_numbers != X0) & (k6_numbers = X0 | ~r1_int_1(k6_numbers,X0))) | ~v1_int_1(X0)) [nnf transformation 92]
137. ! [X0] : (! [X1] : (! [X2] : (((k3_int_2(X0,X1) = X2 | (? [X3] : (~r1_int_1(X3,X2) & r1_int_1(X3,X1) & r1_int_1(X3,X0) & v1_int_1(X3)) | ~r1_int_1(X2,X1) | ~r1_int_1(X2,X0))) & ((! [X3] : (r1_int_1(X3,X2) | ~r1_int_1(X3,X1) | ~r1_int_1(X3,X0) | ~v1_int_1(X3)) & r1_int_1(X2,X1) & r1_int_1(X2,X0)) | k3_int_2(X0,X1) != X2)) | ~v7_ordinal1(X2)) | ~v1_int_1(X1)) | ~v1_int_1(X0)) [nnf transformation 94]
138. ! [X0] : (! [X1] : (! [X2] : (((k3_int_2(X0,X1) = X2 | ? [X3] : (~r1_int_1(X3,X2) & r1_int_1(X3,X1) & r1_int_1(X3,X0) & v1_int_1(X3)) | ~r1_int_1(X2,X1) | ~r1_int_1(X2,X0)) & ((! [X3] : (r1_int_1(X3,X2) | ~r1_int_1(X3,X1) | ~r1_int_1(X3,X0) | ~v1_int_1(X3)) & r1_int_1(X2,X1) & r1_int_1(X2,X0)) | k3_int_2(X0,X1) != X2)) | ~v7_ordinal1(X2)) | ~v1_int_1(X1)) | ~v1_int_1(X0)) [flattening 137]
139. ! [X0] : (! [X1] : (! [X2] : (((k3_int_2(X0,X1) = X2 | ? [X3] : (~r1_int_1(X3,X2) & r1_int_1(X3,X1) & r1_int_1(X3,X0) & v1_int_1(X3)) | ~r1_int_1(X2,X1) | ~r1_int_1(X2,X0)) & ((! [X4] : (r1_int_1(X4,X2) | ~r1_int_1(X4,X1) | ~r1_int_1(X4,X0) | ~v1_int_1(X4)) & r1_int_1(X2,X1) & r1_int_1(X2,X0)) | k3_int_2(X0,X1) != X2)) | ~v7_ordinal1(X2)) | ~v1_int_1(X1)) | ~v1_int_1(X0)) [rectify 138]
140. ! [X2,X1,X0] : (? [X3] : (~r1_int_1(X3,X2) & r1_int_1(X3,X1) & r1_int_1(X3,X0) & v1_int_1(X3)) => (~r1_int_1(sK2(X0,X1,X2),X2) & r1_int_1(sK2(X0,X1,X2),X1) & r1_int_1(sK2(X0,X1,X2),X0) & v1_int_1(sK2(X0,X1,X2)))) [choice axiom]
141. ! [X0] : (! [X1] : (! [X2] : (((k3_int_2(X0,X1) = X2 | (~r1_int_1(sK2(X0,X1,X2),X2) & r1_int_1(sK2(X0,X1,X2),X1) & r1_int_1(sK2(X0,X1,X2),X0) & v1_int_1(sK2(X0,X1,X2))) | ~r1_int_1(X2,X1) | ~r1_int_1(X2,X0)) & ((! [X4] : (r1_int_1(X4,X2) | ~r1_int_1(X4,X1) | ~r1_int_1(X4,X0) | ~v1_int_1(X4)) & r1_int_1(X2,X1) & r1_int_1(X2,X0)) | k3_int_2(X0,X1) != X2)) | ~v7_ordinal1(X2)) | ~v1_int_1(X1)) | ~v1_int_1(X0)) [skolemisation 139,140]
169. v1_int_1(sK0) [cnf transformation 135] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
170. v1_int_1(sK1) [cnf transformation 135] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
171. k6_numbers = k3_int_2(sK0,sK1) | k6_numbers = sK0 [cnf transformation 135] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
172. k6_numbers = k3_int_2(sK0,sK1) | k6_numbers = sK1 [cnf transformation 135] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
173. k6_numbers != k3_int_2(sK0,sK1) | k6_numbers != sK1 | k6_numbers != sK0 [cnf transformation 135] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
175. k5_numbers = k4_ordinal1 [cnf transformation 50] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
176. k6_numbers = k1_xboole_0 [cnf transformation 51] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
177. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 55] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
183. ~r1_int_1(k6_numbers,X0) | k6_numbers = X0 | ~v1_int_1(X0) [cnf transformation 136] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
184. r1_int_1(k6_numbers,X0) | k6_numbers != X0 | ~v1_int_1(X0) [cnf transformation 136] {a:0,w:8,wCS:3200,sine:255,thAx:0,allAx:1,thDist:-1}
185. r1_int_1(X2,X0) | k3_int_2(X0,X1) != X2 | ~v7_ordinal1(X2) | ~v1_int_1(X1) | ~v1_int_1(X0) [cnf transformation 141] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:2,thDist:-2}
186. r1_int_1(X2,X1) | k3_int_2(X0,X1) != X2 | ~v7_ordinal1(X2) | ~v1_int_1(X1) | ~v1_int_1(X0) [cnf transformation 141] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:2,thDist:-2}
189. r1_int_1(sK2(X0,X1,X2),X0) | k3_int_2(X0,X1) = X2 | ~r1_int_1(X2,X1) | ~r1_int_1(X2,X0) | ~v7_ordinal1(X2) | ~v1_int_1(X1) | ~v1_int_1(X0) [cnf transformation 141] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
191. ~r1_int_1(sK2(X0,X1,X2),X2) | k3_int_2(X0,X1) = X2 | ~r1_int_1(X2,X1) | ~r1_int_1(X2,X0) | ~v7_ordinal1(X2) | ~v1_int_1(X1) | ~v1_int_1(X0) [cnf transformation 141] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
192. ~v7_ordinal1(X0) | v1_int_1(X0) [cnf transformation 95] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
206. v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1) [cnf transformation 106] {a:0,w:5,wCS:2000,sine:3,thAx:0,allAx:1,thDist:-1}
223. v7_ordinal1(k3_int_2(X0,X1)) | ~v1_int_1(X1) | ~v1_int_1(X0) [cnf transformation 128] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
243. r1_int_1(k6_numbers,k6_numbers) | ~v1_int_1(k6_numbers) [equality resolution 184] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
245. r1_int_1(k3_int_2(X0,X1),X1) | ~v7_ordinal1(k3_int_2(X0,X1)) | ~v1_int_1(X1) | ~v1_int_1(X0) [equality resolution 186] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:2,thDist:-2}
246. r1_int_1(k3_int_2(X0,X1),X0) | ~v7_ordinal1(k3_int_2(X0,X1)) | ~v1_int_1(X1) | ~v1_int_1(X0) [equality resolution 185] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:2,thDist:-2}
248. 1 <=> k6_numbers = sK0 [avatar definition]
249. k6_numbers = sK0 <- (1) [avatar component clause 248] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
250. k6_numbers != sK0 <- (~1) [avatar component clause 248] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
252. 2 <=> k6_numbers = sK1 [avatar definition]
253. k6_numbers = sK1 <- (2) [avatar component clause 252] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
256. 3 <=> k6_numbers = k3_int_2(sK0,sK1) [avatar definition]
257. k6_numbers = k3_int_2(sK0,sK1) <- (3) [avatar component clause 256] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
258. k6_numbers != k3_int_2(sK0,sK1) <- (~3) [avatar component clause 256] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
259. ~1 | ~2 | ~3 [avatar split clause 173,256,252,248]
260. 2 | 3 [avatar split clause 172,256,252]
261. 1 | 3 [avatar split clause 171,256,248]
263. 4 <=> v1_int_1(sK1) [avatar definition]
265. v1_int_1(sK1) <- (4) [avatar component clause 263] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
266. 4 [avatar split clause 170,263]
268. 5 <=> v1_int_1(sK0) [avatar definition]
270. v1_int_1(sK0) <- (5) [avatar component clause 268] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
271. 5 [avatar split clause 169,268]
278. 7 <=> k5_numbers = k4_ordinal1 [avatar definition]
280. k5_numbers = k4_ordinal1 <- (7) [avatar component clause 278] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
281. 7 [avatar split clause 175,278]
283. 8 <=> k6_numbers = k1_xboole_0 [avatar definition]
285. k6_numbers = k1_xboole_0 <- (8) [avatar component clause 283] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
286. 8 [avatar split clause 176,283]
288. 9 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
290. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (9) [avatar component clause 288] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
291. 9 [avatar split clause 177,288]
313. 14 <=> v1_int_1(k6_numbers) [avatar definition]
314. v1_int_1(k6_numbers) <- (14) [avatar component clause 313] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
317. 15 <=> r1_int_1(k6_numbers,k6_numbers) [avatar definition]
319. r1_int_1(k6_numbers,k6_numbers) <- (15) [avatar component clause 317] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
320. ~14 | 15 [avatar split clause 243,317,313]
442. m1_subset_1(k1_xboole_0,k5_numbers) <- (7, 9) [forward demodulation 290,280] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
443. m1_subset_1(k6_numbers,k5_numbers) <- (7, 8, 9) [forward demodulation 442,285] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
445. 39 <=> m1_subset_1(k6_numbers,k5_numbers) [avatar definition]
447. m1_subset_1(k6_numbers,k5_numbers) <- (39) [avatar component clause 445] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
448. 39 | ~7 | ~8 | ~9 [avatar split clause 443,288,283,278,445]
573. 60 <=> v7_ordinal1(k6_numbers) [avatar definition]
575. v7_ordinal1(k6_numbers) <- (60) [avatar component clause 573] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
608. ~m1_subset_1(X0,k5_numbers) | v7_ordinal1(X0) <- (7) [forward demodulation 206,280] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
614. v7_ordinal1(k6_numbers) <- (7, 39) [resolution 608,447] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
616. 60 | ~7 | ~39 [avatar split clause 614,445,278,573]
662. v1_int_1(k6_numbers) <- (60) [resolution 575,192] {a:2,w:2,wCS:800,sine:1,thAx:0,allAx:3,thDist:-3}
667. 14 | ~60 [avatar split clause 662,573,313]
789. r1_int_1(k3_int_2(X0,X1),X1) | ~v1_int_1(X1) | ~v1_int_1(X0) [subsumption resolution 245,223] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
790. r1_int_1(k6_numbers,sK1) | ~v1_int_1(sK1) | ~v1_int_1(sK0) <- (3) [superposition 789,257] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
795. r1_int_1(k6_numbers,sK1) | ~v1_int_1(sK0) <- (3, 4) [subsumption resolution 790,265] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
796. r1_int_1(k6_numbers,sK1) <- (3, 4, 5) [subsumption resolution 795,270] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
798. 88 <=> r1_int_1(k6_numbers,sK1) [avatar definition]
800. r1_int_1(k6_numbers,sK1) <- (88) [avatar component clause 798] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
801. 88 | ~3 | ~4 | ~5 [avatar split clause 796,268,263,256,798]
802. k6_numbers = sK1 | ~v1_int_1(sK1) <- (88) [resolution 800,183] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
804. k6_numbers = sK1 <- (4, 88) [subsumption resolution 802,265] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
805. 2 | ~4 | ~88 [avatar split clause 804,798,263,252]
807. k6_numbers = k3_int_2(sK0,k6_numbers) <- (2, 3) [backward demodulation 257,253] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
810. 89 <=> k6_numbers = k3_int_2(sK0,k6_numbers) [avatar definition]
811. k6_numbers != k3_int_2(sK0,k6_numbers) <- (~89) [avatar component clause 810] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
812. k6_numbers = k3_int_2(sK0,k6_numbers) <- (89) [avatar component clause 810] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
813. 89 | ~2 | ~3 [avatar split clause 807,256,252,810]
837. r1_int_1(k3_int_2(X0,X1),X0) | ~v1_int_1(X1) | ~v1_int_1(X0) [subsumption resolution 246,223] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
840. r1_int_1(k6_numbers,sK0) | ~v1_int_1(k6_numbers) | ~v1_int_1(sK0) <- (89) [superposition 837,812] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
843. r1_int_1(k6_numbers,sK0) | ~v1_int_1(sK0) <- (14, 89) [subsumption resolution 840,314] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
844. r1_int_1(k6_numbers,sK0) <- (5, 14, 89) [subsumption resolution 843,270] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
846. 91 <=> r1_int_1(k6_numbers,sK0) [avatar definition]
848. r1_int_1(k6_numbers,sK0) <- (91) [avatar component clause 846] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
849. 91 | ~5 | ~14 | ~89 [avatar split clause 844,810,313,268,846]
851. k6_numbers = sK0 | ~v1_int_1(sK0) <- (91) [resolution 848,183] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
854. ~v1_int_1(sK0) <- (~1, 91) [subsumption resolution 851,250] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
855. $false <- (~1, 5, 91) [subsumption resolution 854,270] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
856. 1 | ~5 | ~91 [avatar contradiction clause 855]
858. k6_numbers != k3_int_2(sK0,k6_numbers) <- (2, ~3) [forward demodulation 258,253] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
859. ~89 | ~2 | 3 [avatar split clause 858,256,252,810]
861. k6_numbers != k3_int_2(k6_numbers,k6_numbers) <- (1, ~89) [forward demodulation 811,249] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
863. 92 <=> k6_numbers = k3_int_2(k6_numbers,k6_numbers) [avatar definition]
865. k6_numbers != k3_int_2(k6_numbers,k6_numbers) <- (~92) [avatar component clause 863] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
866. ~92 | ~1 | 89 [avatar split clause 861,810,248,863]
902. k3_int_2(X0,X1) = X0 | ~r1_int_1(X0,X1) | ~r1_int_1(X0,X0) | ~v7_ordinal1(X0) | ~v1_int_1(X1) | ~v1_int_1(X0) | k3_int_2(X0,X1) = X0 | ~r1_int_1(X0,X1) | ~r1_int_1(X0,X0) | ~v7_ordinal1(X0) | ~v1_int_1(X1) | ~v1_int_1(X0) [resolution 191,189] {a:1,w:34,wCS:13600,sine:255,thAx:0,allAx:4,thDist:-4}
908. k3_int_2(X0,X1) = X0 | ~r1_int_1(X0,X1) | ~r1_int_1(X0,X0) | ~v7_ordinal1(X0) | ~v1_int_1(X1) | ~v1_int_1(X0) [duplicate literal removal 902] {a:1,w:17,wCS:6800,sine:255,thAx:0,allAx:4,thDist:-4}
909. k3_int_2(X0,X1) = X0 | ~r1_int_1(X0,X1) | ~r1_int_1(X0,X0) | ~v7_ordinal1(X0) | ~v1_int_1(X1) [subsumption resolution 908,192] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
1152. k6_numbers != k6_numbers | ~r1_int_1(k6_numbers,k6_numbers) | ~r1_int_1(k6_numbers,k6_numbers) | ~v7_ordinal1(k6_numbers) | ~v1_int_1(k6_numbers) <- (~92) [superposition 865,909] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1153. k6_numbers != k6_numbers | ~r1_int_1(k6_numbers,k6_numbers) | ~v7_ordinal1(k6_numbers) | ~v1_int_1(k6_numbers) <- (~92) [duplicate literal removal 1152] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1154. ~r1_int_1(k6_numbers,k6_numbers) | ~v7_ordinal1(k6_numbers) | ~v1_int_1(k6_numbers) <- (~92) [trivial inequality removal 1153] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1175. ~v7_ordinal1(k6_numbers) | ~v1_int_1(k6_numbers) <- (15, ~92) [subsumption resolution 1154,319] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1176. ~v1_int_1(k6_numbers) <- (15, 60, ~92) [subsumption resolution 1175,575] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1177. $false <- (14, 15, 60, ~92) [subsumption resolution 1176,314] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1178. ~14 | ~15 | ~60 | 92 [avatar contradiction clause 1177]
1179. $false [avatar sat refutation 259,260,261,266,271,281,286,291,320,448,616,667,801,805,813,849,856,859,866,1178] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:628,thDist:-628}
% SZS output end Proof for int_2__t5_int_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 63
% >>> Preprocessing
% Introduced skolems: 16
% Pure predicates: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 74
% Generated clauses: 556
% Active clauses: 169
% Passive clauses: 234
% Final active clauses: 155
% Final passive clauses: 48
% Main loop iterations started: 183
% 
% >>> Simplifying Inferences
% Duplicate literals: 69
% Trivial inequalities: 2
% Fw subsumption resolutions: 64
% Fw demodulations: 8
% Bw demodulations: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Equational tautologies: 7
% Forward subsumptions: 69
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 188
% Unit resulting resolution: 4
% Forward superposition: 36
% Backward superposition: 20
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 13
% Split components: 27
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 127
% SAT solver unit clauses: 29
% SAT solver binary clauses: 59
% 
% Memory used [KB]: 5628
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% congruence closure: 0.001 s
% forward subsumption: 0.002 s
% resolution: 0.001 s
% other: 0.008 s ( own 0.002 s ) 

