i: [1,31,0,2,1,1,0,0] t70_relat_1
a: [3,51,0,2,1,31]
s: 51
i: [1,32,0,7,1,1,0,0] t70_relat_1
a: [3,46,0,7,1,32]
s: 46
i: [1,41,0,3,1,0,0,1] reflexivity_r1_tarski
s: 41
i: [1,40,0,4,1,0,0,255] existence_m1_subset_1
s: 40
i: [1,42,0,6,2,0,0,1] dt_k5_relat_1
s: 42
i: [1,37,0,7,2,0,0,1] t59_relat_1
s: 37
i: [1,38,0,7,2,0,0,255] t3_subset
s: 38
i: [1,39,0,7,2,0,0,255] t3_subset
s: 39
d: [2,53,1,5,1,0,40,38,40]
s: 53
i: [1,33,0,8,2,0,0,255] rd4_relat_1
s: 33
i: [1,36,0,8,3,0,0,255] cc2_relat_1
s: 36
d: [2,59,1,6,2,0,40,36,40]
s: 59
i: [1,34,0,12,4,0,0,1] t11_relat_1
s: 34
d: [2,60,1,7,3,0,40,36,39]
s: 60
i: [1,35,0,12,4,0,0,1] t11_relat_1
s: 35
d: [2,67,0,10,3,0,21,35,60]
s: 67
d: [2,68,1,7,2,1,40,67,46]
d: [2,70,1,5,1,2,21,68,51]
a: [3,74,1,5,1,70]
s: 74
d: [2,76,2,0,0,2,52,51,74,37]
e: 76
d: [2,77,2,2,1,1,40,74,37]
d: [2,79,2,0,0,2,21,77,51]
e: 79
f: 79,70,31,32
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t70_relat_1
% SZS output start Proof for relat_1__t70_relat_1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => r1_tarski(k10_xtuple_0(k5_relat_1(X1,X0)),k10_xtuple_0(X1))) [input t70_relat_1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => r1_tarski(k10_xtuple_0(k5_relat_1(X1,X0)),k10_xtuple_0(X1))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
12. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,X1) => (r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)))))) [input t11_relat_1]
13. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
14. ! [X0] : ! [X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [input t59_relat_1]
15. ~! [X0,X1] : (v1_relat_1(X1) => r1_tarski(k10_xtuple_0(k5_relat_1(X1,X0)),k10_xtuple_0(X1))) [flattening 2]
16. ! [X0,X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [flattening 14]
17. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 13]
19. ? [X0,X1] : (~r1_tarski(k10_xtuple_0(k5_relat_1(X1,X0)),k10_xtuple_0(X1)) & v1_relat_1(X1)) [ennf transformation 15]
21. ! [X0] : (! [X1] : (((r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 12]
22. ! [X0] : (! [X1] : ((r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 21]
23. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
24. ! [X0,X1] : (r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1)) [ennf transformation 16]
26. ? [X0,X1] : (~r1_tarski(k10_xtuple_0(k5_relat_1(X1,X0)),k10_xtuple_0(X1)) & v1_relat_1(X1)) => (~r1_tarski(k10_xtuple_0(k5_relat_1(sK1,sK0)),k10_xtuple_0(sK1)) & v1_relat_1(sK1)) [choice axiom]
27. ~r1_tarski(k10_xtuple_0(k5_relat_1(sK1,sK0)),k10_xtuple_0(sK1)) & v1_relat_1(sK1) [skolemisation 19,26]
28. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 17]
31. v1_relat_1(sK1) [cnf transformation 27] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. ~r1_tarski(k10_xtuple_0(k5_relat_1(sK1,sK0)),k10_xtuple_0(sK1)) [cnf transformation 27] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
36. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
37. r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1) [cnf transformation 24] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
39. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 28] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. 1 <=> r1_tarski(k10_xtuple_0(k5_relat_1(sK1,sK0)),k10_xtuple_0(sK1)) [avatar definition]
46. ~r1_tarski(k10_xtuple_0(k5_relat_1(sK1,sK0)),k10_xtuple_0(sK1)) <- (~1) [avatar component clause 44] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
47. ~1 [avatar split clause 32,44]
49. 2 <=> v1_relat_1(sK1) [avatar definition]
51. v1_relat_1(sK1) <- (2) [avatar component clause 49] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
52. 2 [avatar split clause 31,49]
60. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 36,39] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
67. r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) [subsumption resolution 35,60] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. ~r1_tarski(k5_relat_1(sK1,sK0),sK1) | ~v1_relat_1(sK1) <- (~1) [resolution 67,46] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
70. ~r1_tarski(k5_relat_1(sK1,sK0),sK1) <- (~1, 2) [subsumption resolution 68,51] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
72. 3 <=> r1_tarski(k5_relat_1(sK1,sK0),sK1) [avatar definition]
74. ~r1_tarski(k5_relat_1(sK1,sK0),sK1) <- (~3) [avatar component clause 72] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
75. ~3 | 1 | ~2 [avatar split clause 70,49,44,72]
77. ~v1_relat_1(sK1) <- (~3) [resolution 74,37] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
79. $false <- (2, ~3) [subsumption resolution 77,51] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
80. ~2 | 3 [avatar contradiction clause 79]
81. $false [avatar sat refutation 47,52,75,80] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:22,thDist:-22}
% SZS output end Proof for relat_1__t70_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 36
% Active clauses: 16
% Passive clauses: 19
% Final active clauses: 16
% Final passive clauses: 2
% Main loop iterations started: 17
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 1
% Backward superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 7
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

