i: [1,40,0,2,1,1,0,0] t93_funct_1
a: [3,71,0,2,1,40]
s: 71
i: [1,41,0,4,1,1,0,0] t93_funct_1
a: [3,66,0,4,1,41]
s: 66
i: [1,43,0,5,1,1,0,0] t93_funct_1
a: [3,56,0,5,1,43]
s: 56
i: [1,42,0,5,1,1,0,0] t93_funct_1
a: [3,61,0,5,1,42]
s: 61
i: [1,52,0,2,1,0,0,1] rc1_funct_1
a: [3,76,0,2,1,52]
s: 76
i: [1,51,0,3,1,0,0,1] reflexivity_r1_tarski
s: 51
i: [1,50,0,4,1,0,0,255] existence_m1_subset_1
s: 50
i: [1,46,0,7,2,0,0,255] t3_subset
s: 46
d: [2,78,1,5,1,0,40,46,50]
s: 78
i: [1,47,0,7,2,0,0,255] t3_subset
s: 47
i: [1,44,0,8,3,0,0,255] cc2_relat_1
s: 44
d: [2,80,1,6,2,0,40,44,50]
s: 80
i: [1,49,0,9,3,0,0,1] t1_xboole_1
s: 49
d: [2,84,1,7,2,1,40,49,66]
s: 84
i: [1,48,0,12,3,0,0,1] t143_relat_1
s: 48
i: [1,45,0,13,3,0,0,1] t76_funct_1
s: 45
d: [2,83,1,8,2,1,40,49,61]
s: 83
d: [2,81,1,7,3,0,40,44,47]
s: 81
d: [2,99,2,6,2,1,40,81,61]
a: [3,118,2,2,1,99]
s: 118
d: [2,87,1,15,4,0,40,48,49]
s: 87
d: [2,88,1,16,4,0,40,45,49]
s: 88
d: [2,97,2,8,3,1,40,81,84]
a: [3,109,2,3,1,97]
s: 109
d: [2,90,2,7,1,1,40,83,78]
a: [3,94,2,7,1,90]
s: 94
d: [2,86,2,10,3,1,40,84,49]
s: 86
d: [2,85,2,8,2,0,40,49,78]
s: 85
d: [2,100,2,13,4,0,40,81,48]
s: 100
d: [2,136,3,10,2,1,40,94,49]
s: 136
d: [2,98,2,14,4,0,40,81,45]
s: 98
d: [2,142,4,9,1,1,40,136,78]
a: [3,146,4,9,1,142]
s: 146
d: [2,128,2,14,3,0,40,87,78]
s: 128
d: [2,127,2,17,5,0,40,87,48]
d: [2,129,2,15,4,0,18,127]
s: 129
d: [2,149,5,12,2,1,40,146,49]
s: 149
d: [2,125,2,18,5,0,40,87,45]
d: [2,130,2,16,4,0,18,125]
s: 130
d: [2,165,3,0,0,4,52,71,66,56,61,130]
e: 165
d: [2,166,3,11,3,1,40,130,61]
d: [2,173,3,6,2,2,21,166,56]
d: [2,174,3,4,1,3,21,173,71]
d: [2,175,3,0,0,4,21,174,66]
e: 175
f: 175,40,41,42,43
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_1__t93_funct_1
% SZS output start Proof for funct_1__t93_funct_1
1. ! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_tarski(k7_relat_1(X2,X0),X1) & r1_tarski(X0,k9_xtuple_0(X2))) => r1_tarski(X0,k8_relat_1(X2,X1)))) [input t93_funct_1]
2. ~! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_tarski(k7_relat_1(X2,X0),X1) & r1_tarski(X0,k9_xtuple_0(X2))) => r1_tarski(X0,k8_relat_1(X2,X1)))) [negated conjecture 1]
12. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => r1_tarski(k8_relat_1(X2,X0),k8_relat_1(X2,X1)))) [input t143_relat_1]
13. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
15. ! [X0] : ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,k9_xtuple_0(X1)) => r1_tarski(X0,k8_relat_1(X1,k7_relat_1(X1,X0))))) [input t76_funct_1]
16. ~! [X0,X1,X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_tarski(k7_relat_1(X2,X0),X1) & r1_tarski(X0,k9_xtuple_0(X2))) => r1_tarski(X0,k8_relat_1(X2,X1)))) [flattening 2]
17. ! [X0,X1] : (v1_relat_1(X1) => (r1_tarski(X0,k9_xtuple_0(X1)) => r1_tarski(X0,k8_relat_1(X1,k7_relat_1(X1,X0))))) [flattening 15]
19. ! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => r1_tarski(k8_relat_1(X2,X0),k8_relat_1(X2,X1)))) [flattening 12]
20. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 13]
23. ~! [X0,X1,X2] : (v1_relat_1(X2) => ((r1_tarski(k7_relat_1(X2,X0),X1) & r1_tarski(X0,k9_xtuple_0(X2))) => r1_tarski(X0,k8_relat_1(X2,X1)))) [pure predicate removal 16]
24. ? [X0,X1,X2] : ((~r1_tarski(X0,k8_relat_1(X2,X1)) & (r1_tarski(k7_relat_1(X2,X0),X1) & r1_tarski(X0,k9_xtuple_0(X2)))) & v1_relat_1(X2)) [ennf transformation 23]
25. ? [X0,X1,X2] : (~r1_tarski(X0,k8_relat_1(X2,X1)) & r1_tarski(k7_relat_1(X2,X0),X1) & r1_tarski(X0,k9_xtuple_0(X2)) & v1_relat_1(X2)) [flattening 24]
27. ! [X0,X1] : ((r1_tarski(X0,k8_relat_1(X1,k7_relat_1(X1,X0))) | ~r1_tarski(X0,k9_xtuple_0(X1))) | ~v1_relat_1(X1)) [ennf transformation 17]
28. ! [X0,X1] : (r1_tarski(X0,k8_relat_1(X1,k7_relat_1(X1,X0))) | ~r1_tarski(X0,k9_xtuple_0(X1)) | ~v1_relat_1(X1)) [flattening 27]
29. ! [X0,X1,X2] : ((r1_tarski(k8_relat_1(X2,X0),k8_relat_1(X2,X1)) | ~r1_tarski(X0,X1)) | ~v1_relat_1(X2)) [ennf transformation 19]
30. ! [X0,X1,X2] : (r1_tarski(k8_relat_1(X2,X0),k8_relat_1(X2,X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X2)) [flattening 29]
31. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 20]
32. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 31]
33. ? [X0,X1,X2] : (~r1_tarski(X0,k8_relat_1(X2,X1)) & r1_tarski(k7_relat_1(X2,X0),X1) & r1_tarski(X0,k9_xtuple_0(X2)) & v1_relat_1(X2)) => (~r1_tarski(sK0,k8_relat_1(sK2,sK1)) & r1_tarski(k7_relat_1(sK2,sK0),sK1) & r1_tarski(sK0,k9_xtuple_0(sK2)) & v1_relat_1(sK2)) [choice axiom]
34. ~r1_tarski(sK0,k8_relat_1(sK2,sK1)) & r1_tarski(k7_relat_1(sK2,sK0),sK1) & r1_tarski(sK0,k9_xtuple_0(sK2)) & v1_relat_1(sK2) [skolemisation 25,33]
40. v1_relat_1(sK2) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
41. r1_tarski(sK0,k9_xtuple_0(sK2)) [cnf transformation 34] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. r1_tarski(k7_relat_1(sK2,sK0),sK1) [cnf transformation 34] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. ~r1_tarski(sK0,k8_relat_1(sK2,sK1)) [cnf transformation 34] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. r1_tarski(X0,k8_relat_1(X1,k7_relat_1(X1,X0))) | ~r1_tarski(X0,k9_xtuple_0(X1)) | ~v1_relat_1(X1) [cnf transformation 28] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
48. r1_tarski(k8_relat_1(X2,X0),k8_relat_1(X2,X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X2) [cnf transformation 30] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 32] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. 1 <=> r1_tarski(sK0,k8_relat_1(sK2,sK1)) [avatar definition]
56. ~r1_tarski(sK0,k8_relat_1(sK2,sK1)) <- (~1) [avatar component clause 54] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
57. ~1 [avatar split clause 43,54]
59. 2 <=> r1_tarski(k7_relat_1(sK2,sK0),sK1) [avatar definition]
61. r1_tarski(k7_relat_1(sK2,sK0),sK1) <- (2) [avatar component clause 59] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. 2 [avatar split clause 42,59]
64. 3 <=> r1_tarski(sK0,k9_xtuple_0(sK2)) [avatar definition]
66. r1_tarski(sK0,k9_xtuple_0(sK2)) <- (3) [avatar component clause 64] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
67. 3 [avatar split clause 41,64]
69. 4 <=> v1_relat_1(sK2) [avatar definition]
71. v1_relat_1(sK2) <- (4) [avatar component clause 69] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. 4 [avatar split clause 40,69]
87. ~r1_tarski(X3,k8_relat_1(X2,X0)) | ~v1_relat_1(X2) | r1_tarski(X3,k8_relat_1(X2,X1)) | ~r1_tarski(X0,X1) [resolution 48,49] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
125. ~v1_relat_1(X0) | r1_tarski(X1,k8_relat_1(X0,X2)) | ~r1_tarski(k7_relat_1(X0,X1),X2) | ~r1_tarski(X1,k9_xtuple_0(X0)) | ~v1_relat_1(X0) [resolution 87,45] {a:2,w:18,wCS:7200,sine:1,thAx:0,allAx:3,thDist:-3}
130. ~r1_tarski(k7_relat_1(X0,X1),X2) | r1_tarski(X1,k8_relat_1(X0,X2)) | ~v1_relat_1(X0) | ~r1_tarski(X1,k9_xtuple_0(X0)) [duplicate literal removal 125] {a:2,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
166. r1_tarski(sK0,k8_relat_1(sK2,sK1)) | ~v1_relat_1(sK2) | ~r1_tarski(sK0,k9_xtuple_0(sK2)) <- (2) [resolution 130,61] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
173. ~v1_relat_1(sK2) | ~r1_tarski(sK0,k9_xtuple_0(sK2)) <- (~1, 2) [subsumption resolution 166,56] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
174. ~r1_tarski(sK0,k9_xtuple_0(sK2)) <- (~1, 2, 4) [subsumption resolution 173,71] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
175. $false <- (~1, 2, 3, 4) [subsumption resolution 174,66] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
176. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 175]
177. $false [avatar sat refutation 57,62,67,72,176] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:89,thDist:-89}
% SZS output end Proof for funct_1__t93_funct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 83
% Active clauses: 33
% Passive clauses: 45
% Final active clauses: 33
% Final passive clauses: 12
% Main loop iterations started: 33
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Fw subsumption resolutions: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 51
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 5
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

