i: [1,38,0,2,1,1,0,0] t42_cgames_1
a: [3,63,0,2,1,38]
s: 63
i: [1,37,0,2,1,1,0,0] t42_cgames_1
a: [3,68,0,2,1,37]
s: 68
i: [1,36,0,2,1,1,0,0] t42_cgames_1
a: [3,73,0,2,1,36]
s: 73
i: [1,35,0,2,1,1,0,0] t42_cgames_1
a: [3,78,0,2,1,35]
s: 78
i: [1,34,0,2,1,1,0,0] t42_cgames_1
a: [3,83,0,2,1,34]
s: 83
i: [1,59,0,2,1,0,0,1] rc4_cgames_1
a: [3,88,0,2,1,59]
s: 88
i: [1,42,0,6,3,0,0,1] cc11_cgames_1
s: 42
i: [1,43,0,6,3,0,0,1] cc11_cgames_1
s: 43
i: [1,45,0,6,3,0,0,1] cc9_cgames_1
s: 45
i: [1,46,0,6,3,0,0,1] cc9_cgames_1
s: 46
i: [1,48,0,6,3,0,0,1] cc7_cgames_1
s: 48
i: [1,49,0,6,3,0,0,1] cc7_cgames_1
s: 49
i: [1,57,0,6,3,0,0,2] cc5_cgames_1
s: 57
i: [1,58,0,6,3,0,0,2] cc5_cgames_1
s: 58
i: [1,40,0,8,4,0,0,1] cc12_cgames_1
s: 40
i: [1,51,0,8,4,0,0,1] cc10_cgames_1
s: 51
d: [2,90,1,6,3,1,40,40,83]
d: [2,92,1,4,2,2,21,90,63]
a: [3,100,1,2,1,92]
s: 100
i: [1,53,0,8,4,0,0,1] cc8_cgames_1
s: 53
d: [2,116,2,0,0,4,52,83,68,78,100,53]
e: 116
d: [2,117,2,6,3,1,40,53,100]
d: [2,119,2,4,2,2,21,117,78]
d: [2,120,2,2,1,3,21,119,68]
d: [2,121,2,0,0,4,21,120,83]
e: 121
r: 100
a: [3,96,1,2,1,92]
s: 96
d: [2,123,2,0,0,4,52,83,73,78,96,51]
e: 123
d: [2,124,2,6,3,1,40,96,51]
d: [2,127,2,4,2,2,21,124,78]
d: [2,128,2,2,1,3,21,127,73]
d: [2,129,2,0,0,4,21,128,83]
e: 129
f: 129,121,92,34,35,36,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for cgames_1__t42_cgames_1
% SZS output start Proof for cgames_1__t42_cgames_1
1. ! [X0] : (v2_cgames_1(X0) => ~(~v8_cgames_1(X0) & ~v10_cgames_1(X0) & ~v9_cgames_1(X0) & ~v7_cgames_1(X0))) [input t42_cgames_1]
2. ~! [X0] : (v2_cgames_1(X0) => ~(~v8_cgames_1(X0) & ~v10_cgames_1(X0) & ~v9_cgames_1(X0) & ~v7_cgames_1(X0))) [negated conjecture 1]
3. ! [X0] : ((~v7_cgames_1(X0) & v5_cgames_1(X0) & v2_cgames_1(X0)) => (v9_cgames_1(X0) & v2_cgames_1(X0))) [input cc10_cgames_1]
5. ! [X0] : ((~v6_cgames_1(X0) & ~v5_cgames_1(X0) & v2_cgames_1(X0)) => (v8_cgames_1(X0) & v2_cgames_1(X0))) [input cc12_cgames_1]
9. ! [X0] : ((~v7_cgames_1(X0) & v6_cgames_1(X0) & v2_cgames_1(X0)) => (v10_cgames_1(X0) & v2_cgames_1(X0))) [input cc8_cgames_1]
12. ? [X0] : ((~v8_cgames_1(X0) & ~v10_cgames_1(X0) & ~v9_cgames_1(X0) & ~v7_cgames_1(X0)) & v2_cgames_1(X0)) [ennf transformation 2]
13. ? [X0] : (~v8_cgames_1(X0) & ~v10_cgames_1(X0) & ~v9_cgames_1(X0) & ~v7_cgames_1(X0) & v2_cgames_1(X0)) [flattening 12]
14. ! [X0] : ((v8_cgames_1(X0) & v2_cgames_1(X0)) | (v6_cgames_1(X0) | v5_cgames_1(X0) | ~v2_cgames_1(X0))) [ennf transformation 5]
15. ! [X0] : ((v8_cgames_1(X0) & v2_cgames_1(X0)) | v6_cgames_1(X0) | v5_cgames_1(X0) | ~v2_cgames_1(X0)) [flattening 14]
22. ! [X0] : ((v9_cgames_1(X0) & v2_cgames_1(X0)) | (v7_cgames_1(X0) | ~v5_cgames_1(X0) | ~v2_cgames_1(X0))) [ennf transformation 3]
23. ! [X0] : ((v9_cgames_1(X0) & v2_cgames_1(X0)) | v7_cgames_1(X0) | ~v5_cgames_1(X0) | ~v2_cgames_1(X0)) [flattening 22]
24. ! [X0] : ((v10_cgames_1(X0) & v2_cgames_1(X0)) | (v7_cgames_1(X0) | ~v6_cgames_1(X0) | ~v2_cgames_1(X0))) [ennf transformation 9]
25. ! [X0] : ((v10_cgames_1(X0) & v2_cgames_1(X0)) | v7_cgames_1(X0) | ~v6_cgames_1(X0) | ~v2_cgames_1(X0)) [flattening 24]
30. ? [X0] : (~v8_cgames_1(X0) & ~v10_cgames_1(X0) & ~v9_cgames_1(X0) & ~v7_cgames_1(X0) & v2_cgames_1(X0)) => (~v8_cgames_1(sK0) & ~v10_cgames_1(sK0) & ~v9_cgames_1(sK0) & ~v7_cgames_1(sK0) & v2_cgames_1(sK0)) [choice axiom]
31. ~v8_cgames_1(sK0) & ~v10_cgames_1(sK0) & ~v9_cgames_1(sK0) & ~v7_cgames_1(sK0) & v2_cgames_1(sK0) [skolemisation 13,30]
34. v2_cgames_1(sK0) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. ~v7_cgames_1(sK0) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. ~v9_cgames_1(sK0) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. ~v10_cgames_1(sK0) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. ~v8_cgames_1(sK0) [cnf transformation 31] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. ~v2_cgames_1(X0) | v6_cgames_1(X0) | v5_cgames_1(X0) | v8_cgames_1(X0) [cnf transformation 15] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. ~v5_cgames_1(X0) | v7_cgames_1(X0) | v9_cgames_1(X0) | ~v2_cgames_1(X0) [cnf transformation 23] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. ~v6_cgames_1(X0) | v7_cgames_1(X0) | v10_cgames_1(X0) | ~v2_cgames_1(X0) [cnf transformation 25] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. 1 <=> v8_cgames_1(sK0) [avatar definition]
63. ~v8_cgames_1(sK0) <- (~1) [avatar component clause 61] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
64. ~1 [avatar split clause 38,61]
66. 2 <=> v10_cgames_1(sK0) [avatar definition]
68. ~v10_cgames_1(sK0) <- (~2) [avatar component clause 66] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
69. ~2 [avatar split clause 37,66]
71. 3 <=> v9_cgames_1(sK0) [avatar definition]
73. ~v9_cgames_1(sK0) <- (~3) [avatar component clause 71] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
74. ~3 [avatar split clause 36,71]
76. 4 <=> v7_cgames_1(sK0) [avatar definition]
78. ~v7_cgames_1(sK0) <- (~4) [avatar component clause 76] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. ~4 [avatar split clause 35,76]
81. 5 <=> v2_cgames_1(sK0) [avatar definition]
83. v2_cgames_1(sK0) <- (5) [avatar component clause 81] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. 5 [avatar split clause 34,81]
90. v6_cgames_1(sK0) | v5_cgames_1(sK0) | v8_cgames_1(sK0) <- (5) [resolution 40,83] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
92. v6_cgames_1(sK0) | v5_cgames_1(sK0) <- (~1, 5) [subsumption resolution 90,63] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
94. 7 <=> v5_cgames_1(sK0) [avatar definition]
96. v5_cgames_1(sK0) <- (7) [avatar component clause 94] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. 8 <=> v6_cgames_1(sK0) [avatar definition]
100. v6_cgames_1(sK0) <- (8) [avatar component clause 98] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. 7 | 8 | 1 | ~5 [avatar split clause 92,81,61,98,94]
117. v7_cgames_1(sK0) | v10_cgames_1(sK0) | ~v2_cgames_1(sK0) <- (8) [resolution 53,100] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
119. v10_cgames_1(sK0) | ~v2_cgames_1(sK0) <- (~4, 8) [subsumption resolution 117,78] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
120. ~v2_cgames_1(sK0) <- (~2, ~4, 8) [subsumption resolution 119,68] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
121. $false <- (~2, ~4, 5, 8) [subsumption resolution 120,83] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
122. 2 | 4 | ~5 | ~8 [avatar contradiction clause 121]
124. v7_cgames_1(sK0) | v9_cgames_1(sK0) | ~v2_cgames_1(sK0) <- (7) [resolution 96,51] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
127. v9_cgames_1(sK0) | ~v2_cgames_1(sK0) <- (~4, 7) [subsumption resolution 124,78] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
128. ~v2_cgames_1(sK0) <- (~3, ~4, 7) [subsumption resolution 127,73] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
129. $false <- (~3, ~4, 5, 7) [subsumption resolution 128,83] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
130. 3 | 4 | ~5 | ~7 [avatar contradiction clause 129]
131. $false [avatar sat refutation 64,69,74,79,84,101,122,130] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:60,thDist:-60}
% SZS output end Proof for cgames_1__t42_cgames_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 51
% Active clauses: 19
% Passive clauses: 22
% Final active clauses: 18
% Final passive clauses: 3
% Main loop iterations started: 19
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Forward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 6
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 5
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 6
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

