i: [1,59,0,3,1,1,0,0] t53_zfmisc_1
a: [3,105,0,3,1,59]
s: 105
i: [1,58,0,3,1,1,0,0] t53_zfmisc_1
a: [3,110,0,3,1,58]
s: 110
i: [1,57,0,8,1,1,0,0] t53_zfmisc_1
a: [3,115,0,8,1,57]
s: 115
i: [1,60,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,120,0,2,1,60]
s: 120
i: [1,61,0,2,1,0,0,255] fc1_xboole_0
a: [3,125,0,2,1,61]
s: 125
i: [1,89,0,2,1,0,0,2] rc2_xboole_0
a: [3,135,0,2,1,89]
s: 135
i: [1,90,0,2,1,0,0,2] rc1_xboole_0
a: [3,140,0,2,1,90]
s: 140
i: [1,63,0,3,1,0,0,1] fc2_xboole_0
s: 63
i: [1,62,0,3,1,0,0,255] d2_xboole_0
a: [3,130,0,3,1,62]
s: 130
r: 120
i: [1,84,0,4,1,0,0,1] fc3_xboole_0
s: 84
i: [1,92,0,4,1,0,0,1] d1_tarski
s: 92
i: [1,64,0,5,1,0,0,255] t2_boole
s: 64
i: [1,65,0,5,2,0,0,255] t6_boole
s: 65
i: [1,71,0,5,2,0,0,1] t7_boole
s: 71
i: [1,85,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 85
i: [1,95,0,5,1,0,0,1] d2_tarski
s: 95
d: [2,146,1,2,1,1,40,71,110]
a: [3,151,1,2,1,146]
s: 151
i: [1,97,0,5,1,0,0,1] d2_tarski
s: 97
i: [1,88,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 88
i: [1,70,0,7,3,0,0,2] t8_boole
s: 70
i: [1,86,0,7,1,0,0,1] commutativity_k2_tarski
s: 86
d: [2,155,1,3,1,1,40,88,110]
a: [3,162,1,3,1,155]
s: 162
i: [1,87,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 87
d: [2,170,1,8,1,1,44,87,115]
a: [3,181,1,8,1,170]
s: 181
i: [1,93,0,7,2,0,0,1] d1_tarski
s: 93
i: [1,100,0,8,2,0,0,1] d4_xboole_0
s: 100
d: [2,187,1,7,2,1,44,100,115]
s: 187
i: [1,101,0,8,2,0,0,1] d4_xboole_0
s: 101
d: [2,193,2,3,1,1,40,187,92]
a: [3,197,2,3,1,193]
s: 197
i: [1,98,0,11,3,0,0,1] d2_tarski
s: 98
i: [1,99,0,11,3,0,0,1] d4_xboole_0
s: 99
d: [2,205,3,3,1,1,40,197,88]
a: [3,210,3,3,1,205]
s: 210
i: [1,68,0,14,3,0,0,1] d1_tarski
s: 68
d: [2,199,1,9,2,1,44,101,115]
s: 199
i: [1,69,0,14,3,0,0,1] d1_tarski
s: 69
i: [1,76,0,17,3,0,0,1] d2_tarski
s: 76
d: [2,192,2,9,2,1,44,100,181]
s: 192
i: [1,77,0,17,3,0,0,1] d2_tarski
s: 77
d: [2,235,2,9,2,1,40,199,88]
s: 235
i: [1,81,0,17,3,0,0,1] d4_xboole_0
s: 81
i: [1,82,0,17,3,0,0,1] d4_xboole_0
s: 82
d: [2,240,3,12,2,1,40,235,199]
d: [2,242,3,6,1,1,18,240]
a: [3,251,3,6,1,242]
s: 251
i: [1,75,0,23,4,0,0,1] d2_tarski
s: 75
d: [2,237,3,7,1,1,40,235,92]
a: [3,246,3,7,1,237]
s: 246
i: [1,83,0,23,4,0,0,1] d4_xboole_0
s: 83
d: [2,220,1,12,3,1,44,99,115]
s: 220
d: [2,283,2,7,2,1,40,220,95]
d: [2,289,2,4,1,2,21,283,110]
a: [3,293,2,4,1,289]
s: 293
d: [2,295,3,0,0,2,52,105,293,93]
e: 295
d: [2,297,3,3,1,1,40,293,93]
d: [2,301,3,0,0,2,21,297,105]
e: 301
f: 301,289,57,58,59
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t53_zfmisc_1
% SZS output start Proof for zfmisc_1__t53_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ~(X0 != X1 & r2_hidden(X1,X2) & k3_xboole_0(k2_tarski(X0,X1),X2) = k1_tarski(X0)) [input t53_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ~(X0 != X1 & r2_hidden(X1,X2) & k3_xboole_0(k2_tarski(X0,X1),X2) = k1_tarski(X0)) [negated conjecture 1]
6. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
7. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
9. ! [X0] : ! [X1] : ! [X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [input d4_xboole_0]
25. ~! [X0,X1,X2] : ~(X0 != X1 & r2_hidden(X1,X2) & k3_xboole_0(k2_tarski(X0,X1),X2) = k1_tarski(X0)) [flattening 2]
26. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 6]
29. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 7]
30. ! [X0,X1,X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [flattening 9]
32. ? [X0,X1,X2] : (X0 != X1 & r2_hidden(X1,X2) & k3_xboole_0(k2_tarski(X0,X1),X2) = k1_tarski(X0)) [ennf transformation 25]
37. ? [X0,X1,X2] : (X0 != X1 & r2_hidden(X1,X2) & k3_xboole_0(k2_tarski(X0,X1),X2) = k1_tarski(X0)) => (sK0 != sK1 & r2_hidden(sK1,sK2) & k3_xboole_0(k2_tarski(sK0,sK1),sK2) = k1_tarski(sK0)) [choice axiom]
38. sK0 != sK1 & r2_hidden(sK1,sK2) & k3_xboole_0(k2_tarski(sK0,sK1),sK2) = k1_tarski(sK0) [skolemisation 32,37]
39. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 26]
40. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 39]
41. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 40,41]
43. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 29]
44. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 43]
45. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 44]
46. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
47. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 45,46]
48. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) | ~r2_hidden(X3,X0))) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [nnf transformation 30]
49. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | ~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [flattening 48]
50. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [rectify 49]
51. ! [X2,X1,X0] : (? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2))) => ((~r2_hidden(sK5(X0,X1,X2),X1) | ~r2_hidden(sK5(X0,X1,X2),X0) | ~r2_hidden(sK5(X0,X1,X2),X2)) & ((r2_hidden(sK5(X0,X1,X2),X1) & r2_hidden(sK5(X0,X1,X2),X0)) | r2_hidden(sK5(X0,X1,X2),X2)))) [choice axiom]
52. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ((~r2_hidden(sK5(X0,X1,X2),X1) | ~r2_hidden(sK5(X0,X1,X2),X0) | ~r2_hidden(sK5(X0,X1,X2),X2)) & ((r2_hidden(sK5(X0,X1,X2),X1) & r2_hidden(sK5(X0,X1,X2),X0)) | r2_hidden(sK5(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [skolemisation 50,51]
57. k3_xboole_0(k2_tarski(sK0,sK1),sK2) = k1_tarski(sK0) [cnf transformation 38] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. r2_hidden(sK1,sK2) [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. sK0 != sK1 [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
66. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 42] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
74. r2_hidden(X4,X2) | X1 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 47] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
80. r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) | k3_xboole_0(X0,X1) != X2 [cnf transformation 52] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
93. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 66] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
94. r2_hidden(X4,X2) | k2_tarski(X0,X4) != X2 [equality resolution 74] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
95. r2_hidden(X4,k2_tarski(X0,X4)) [equality resolution 94] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
99. r2_hidden(X4,k3_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) [equality resolution 80] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
103. 1 <=> sK0 = sK1 [avatar definition]
105. sK0 != sK1 <- (~1) [avatar component clause 103] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
106. ~1 [avatar split clause 59,103]
108. 2 <=> r2_hidden(sK1,sK2) [avatar definition]
110. r2_hidden(sK1,sK2) <- (2) [avatar component clause 108] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
111. 2 [avatar split clause 58,108]
113. 3 <=> k3_xboole_0(k2_tarski(sK0,sK1),sK2) = k1_tarski(sK0) [avatar definition]
115. k3_xboole_0(k2_tarski(sK0,sK1),sK2) = k1_tarski(sK0) <- (3) [avatar component clause 113] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. 3 [avatar split clause 57,113]
220. ~r2_hidden(X0,k2_tarski(sK0,sK1)) | ~r2_hidden(X0,sK2) | r2_hidden(X0,k1_tarski(sK0)) <- (3) [superposition 99,115] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
283. ~r2_hidden(sK1,sK2) | r2_hidden(sK1,k1_tarski(sK0)) <- (3) [resolution 220,95] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
289. r2_hidden(sK1,k1_tarski(sK0)) <- (2, 3) [subsumption resolution 283,110] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
291. 16 <=> r2_hidden(sK1,k1_tarski(sK0)) [avatar definition]
293. r2_hidden(sK1,k1_tarski(sK0)) <- (16) [avatar component clause 291] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
294. 16 | ~2 | ~3 [avatar split clause 289,113,108,291]
297. sK0 = sK1 <- (16) [resolution 293,93] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
301. $false <- (~1, 16) [subsumption resolution 297,105] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
302. 1 | ~16 [avatar contradiction clause 301]
303. $false [avatar sat refutation 106,111,116,294,302] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:115,thDist:-115}
% SZS output end Proof for zfmisc_1__t53_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 34
% Generated clauses: 166
% Active clauses: 46
% Passive clauses: 101
% Final active clauses: 45
% Final passive clauses: 54
% Main loop iterations started: 47
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Equational tautologies: 3
% Forward subsumptions: 24
% 
% >>> Generating Inferences
% Binary resolution: 69
% Unit resulting resolution: 1
% Factoring: 2
% Forward superposition: 24
% Backward superposition: 12
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 8
% SAT solver binary clauses: 12
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% superposition: 0.001 s
% literal selection: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

