i: [1,16,0,2,1,1,0,0] l70_xcmplx_1
a: [3,35,0,2,1,16]
s: 35
i: [1,15,0,2,1,1,0,0] l70_xcmplx_1
a: [3,40,0,2,1,15]
s: 40
i: [1,18,0,3,1,1,0,0] l70_xcmplx_1
a: [3,25,0,3,1,18]
s: 25
i: [1,17,0,5,1,1,0,0] l70_xcmplx_1
a: [3,30,0,5,1,17]
s: 30
i: [1,21,0,2,1,0,0,1] rc1_xcmplx_0
a: [3,45,0,2,1,21]
s: 45
i: [1,19,0,5,2,0,0,1] dt_k5_xcmplx_0
s: 19
d: [2,47,1,5,2,1,44,19,30]
d: [2,48,1,3,1,2,21,47,35]
a: [3,52,1,3,1,48]
s: 52
i: [1,20,0,7,2,0,0,1] involutiveness_k5_xcmplx_0
s: 20
d: [2,54,1,7,2,1,44,20,30]
d: [2,56,1,5,1,2,21,54,35]
a: [3,60,1,5,1,56]
s: 60
d: [2,62,2,5,2,1,44,60,20]
d: [2,66,2,2,1,2,21,62,25]
d: [2,67,2,0,0,3,21,66,40]
e: 67
d: [2,63,2,5,2,1,44,20,60]
d: [2,69,2,2,1,2,21,63,25]
d: [2,70,2,0,0,3,21,69,40]
e: 70
f: 70,56,15,16,17,18
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xcmplx_1__l70_xcmplx_1
% SZS output start Proof for xcmplx_1__l70_xcmplx_1
1. ! [X0] : (v1_xcmplx_0(X0) => ! [X1] : (v1_xcmplx_0(X1) => (k5_xcmplx_0(X0) = k5_xcmplx_0(X1) => X0 = X1))) [input l70_xcmplx_1]
2. ~! [X0] : (v1_xcmplx_0(X0) => ! [X1] : (v1_xcmplx_0(X1) => (k5_xcmplx_0(X0) = k5_xcmplx_0(X1) => X0 = X1))) [negated conjecture 1]
4. ! [X0] : (v1_xcmplx_0(X0) => k5_xcmplx_0(k5_xcmplx_0(X0)) = X0) [input involutiveness_k5_xcmplx_0]
6. ? [X0] : (? [X1] : ((X0 != X1 & k5_xcmplx_0(X0) = k5_xcmplx_0(X1)) & v1_xcmplx_0(X1)) & v1_xcmplx_0(X0)) [ennf transformation 2]
7. ? [X0] : (? [X1] : (X0 != X1 & k5_xcmplx_0(X0) = k5_xcmplx_0(X1) & v1_xcmplx_0(X1)) & v1_xcmplx_0(X0)) [flattening 6]
9. ! [X0] : (k5_xcmplx_0(k5_xcmplx_0(X0)) = X0 | ~v1_xcmplx_0(X0)) [ennf transformation 4]
10. ? [X0] : (? [X1] : (X0 != X1 & k5_xcmplx_0(X0) = k5_xcmplx_0(X1) & v1_xcmplx_0(X1)) & v1_xcmplx_0(X0)) => (? [X1] : (sK0 != X1 & k5_xcmplx_0(X1) = k5_xcmplx_0(sK0) & v1_xcmplx_0(X1)) & v1_xcmplx_0(sK0)) [choice axiom]
11. ? [X1] : (sK0 != X1 & k5_xcmplx_0(X1) = k5_xcmplx_0(sK0) & v1_xcmplx_0(X1)) => (sK0 != sK1 & k5_xcmplx_0(sK0) = k5_xcmplx_0(sK1) & v1_xcmplx_0(sK1)) [choice axiom]
12. (sK0 != sK1 & k5_xcmplx_0(sK0) = k5_xcmplx_0(sK1) & v1_xcmplx_0(sK1)) & v1_xcmplx_0(sK0) [skolemisation 7,11,10]
15. v1_xcmplx_0(sK0) [cnf transformation 12] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
16. v1_xcmplx_0(sK1) [cnf transformation 12] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
17. k5_xcmplx_0(sK0) = k5_xcmplx_0(sK1) [cnf transformation 12] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
18. sK0 != sK1 [cnf transformation 12] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
20. k5_xcmplx_0(k5_xcmplx_0(X0)) = X0 | ~v1_xcmplx_0(X0) [cnf transformation 9] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
23. 1 <=> sK0 = sK1 [avatar definition]
25. sK0 != sK1 <- (~1) [avatar component clause 23] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
26. ~1 [avatar split clause 18,23]
28. 2 <=> k5_xcmplx_0(sK0) = k5_xcmplx_0(sK1) [avatar definition]
30. k5_xcmplx_0(sK0) = k5_xcmplx_0(sK1) <- (2) [avatar component clause 28] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
31. 2 [avatar split clause 17,28]
33. 3 <=> v1_xcmplx_0(sK1) [avatar definition]
35. v1_xcmplx_0(sK1) <- (3) [avatar component clause 33] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
36. 3 [avatar split clause 16,33]
38. 4 <=> v1_xcmplx_0(sK0) [avatar definition]
40. v1_xcmplx_0(sK0) <- (4) [avatar component clause 38] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
41. 4 [avatar split clause 15,38]
54. sK1 = k5_xcmplx_0(k5_xcmplx_0(sK0)) | ~v1_xcmplx_0(sK1) <- (2) [superposition 20,30] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
56. sK1 = k5_xcmplx_0(k5_xcmplx_0(sK0)) <- (2, 3) [subsumption resolution 54,35] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
58. 7 <=> sK1 = k5_xcmplx_0(k5_xcmplx_0(sK0)) [avatar definition]
60. sK1 = k5_xcmplx_0(k5_xcmplx_0(sK0)) <- (7) [avatar component clause 58] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
61. 7 | ~2 | ~3 [avatar split clause 56,33,28,58]
63. sK0 = sK1 | ~v1_xcmplx_0(sK0) <- (7) [superposition 20,60] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
69. ~v1_xcmplx_0(sK0) <- (~1, 7) [subsumption resolution 63,25] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
70. $false <- (~1, 4, 7) [subsumption resolution 69,40] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
71. 1 | ~4 | ~7 [avatar contradiction clause 70]
72. $false [avatar sat refutation 26,31,36,41,61,71] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:43,thDist:-43}
% SZS output end Proof for xcmplx_1__l70_xcmplx_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 4
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 7
% Generated clauses: 27
% Active clauses: 9
% Passive clauses: 9
% Final active clauses: 9
% Main loop iterations started: 9
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 2
% 
% >>> Generating Inferences
% Forward superposition: 3
% Backward superposition: 4
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

