i: [1,89,0,5,1,1,0,0] t107_relat_1
a: [3,135,0,5,1,89]
s: 135
i: [1,90,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,140,0,2,1,90]
s: 140
i: [1,91,0,2,1,0,0,1] fc1_xboole_0
a: [3,145,0,2,1,91]
s: 145
i: [1,122,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,155,0,2,1,122]
s: 155
i: [1,123,0,2,1,0,0,2] rc2_xboole_0
a: [3,160,0,2,1,123]
s: 160
i: [1,124,0,2,1,0,0,2] rc1_xboole_0
a: [3,165,0,2,1,124]
s: 165
i: [1,126,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,170,0,2,1,126]
s: 170
i: [1,125,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,175,0,2,1,125]
s: 175
i: [1,128,0,2,1,0,0,2] rc1_relat_1
a: [3,180,0,2,1,128]
s: 180
i: [1,127,0,2,1,0,0,2] rc1_relat_1
a: [3,185,0,2,1,127]
s: 185
i: [1,93,0,3,1,0,0,255] fc2_xboole_0
s: 93
i: [1,94,0,3,1,0,0,255] fc2_zfmisc_1
s: 94
i: [1,92,0,3,1,0,0,255] d2_xboole_0
a: [3,150,0,3,1,92]
s: 150
r: 140
i: [1,95,0,4,2,0,0,255] cc2_zfmisc_1
s: 95
i: [1,96,0,4,2,0,0,255] cc1_zfmisc_1
s: 96
i: [1,97,0,4,2,0,0,2] cc1_relat_1
s: 97
d: [2,188,1,2,1,1,40,95,165]
a: [3,193,1,2,1,188]
s: 193
i: [1,111,0,4,1,0,0,255] existence_m1_subset_1
s: 111
d: [2,189,1,2,1,1,40,95,145]
a: [3,198,1,2,1,189]
s: 198
i: [1,112,0,4,1,0,0,255] fc3_xboole_0
s: 112
i: [1,113,0,4,1,0,0,2] fc1_zfmisc_1
s: 113
d: [2,200,1,2,1,1,40,97,165]
a: [3,205,1,2,1,200]
s: 205
i: [1,98,0,5,2,0,0,1] t6_boole
s: 98
d: [2,215,1,7,2,1,44,135,98]
s: 215
i: [1,110,0,5,2,0,0,2] t7_boole
s: 110
i: [1,114,0,5,1,0,0,255] fc5_relat_1
s: 114
d: [2,201,1,2,1,1,40,97,145]
a: [3,210,1,2,1,201]
s: 210
i: [1,108,0,6,2,0,0,255] t1_subset
s: 108
i: [1,116,0,6,2,0,0,1] dt_k6_relat_1
s: 116
i: [1,118,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 118
i: [1,109,0,7,3,0,0,1] t8_boole
s: 109
i: [1,115,0,7,1,0,0,255] commutativity_k2_tarski
s: 115
i: [1,107,0,8,3,0,0,255] t2_subset
s: 107
d: [2,219,1,6,2,0,40,107,111]
s: 219
i: [1,119,0,8,3,0,0,1] fc18_relat_1
s: 119
i: [1,120,0,8,3,0,0,1] fc18_relat_1
s: 120
d: [2,220,2,6,2,0,40,219,118]
s: 220
i: [1,121,0,8,1,0,0,255] fc7_relat_1
s: 121
d: [2,212,1,7,3,0,44,98,98]
s: 212
i: [1,100,0,10,1,0,0,255] d5_tarski
s: 100
i: [1,117,0,11,2,0,0,1] rd7_relat_1
s: 117
d: [2,236,2,13,3,1,44,215,117]
d: [2,239,2,6,2,1,19,236]
s: 239
i: [1,99,0,12,3,0,0,1] t37_relat_1
s: 99
d: [2,240,3,6,3,1,40,239,119]
d: [2,242,3,4,2,1,18,240]
a: [3,246,3,2,1,242]
s: 246
i: [1,131,0,16,4,0,0,1] d12_relat_1
s: 131
d: [2,254,0,12,3,0,21,131,116]
s: 254
d: [2,224,1,8,3,0,40,119,95]
s: 224
i: [1,130,0,18,4,0,0,1] d12_relat_1
s: 130
d: [2,227,1,10,1,0,44,100,115]
s: 227
d: [2,259,0,14,3,0,21,130,116]
s: 259
i: [1,129,0,21,5,0,0,1] d12_relat_1
s: 129
d: [2,228,1,10,1,0,44,100,115]
s: 228
d: [2,273,0,17,4,0,21,129,116]
s: 273
d: [2,230,1,10,1,0,44,100,115]
s: 230
i: [1,104,0,26,5,0,0,1] d12_relat_1
s: 104
i: [1,105,0,31,5,0,0,1] d12_relat_1
s: 105
d: [2,231,1,10,1,0,44,100,115]
s: 231
i: [1,106,0,37,6,0,0,1] d12_relat_1
s: 106
d: [2,232,1,10,1,0,44,115,100]
s: 232
d: [2,233,1,10,1,0,44,115,100]
s: 233
d: [2,251,1,12,3,0,40,99,118]
s: 251
d: [2,264,2,10,1,0,44,227,115]
s: 264
d: [2,252,1,12,3,0,40,99,108]
s: 252
d: [2,265,2,10,1,0,44,227,115]
s: 265
d: [2,255,1,17,4,0,40,254,99]
d: [2,257,1,13,3,0,21,255,116]
s: 257
d: [2,278,1,14,4,0,40,273,110]
s: 278
d: [2,327,2,9,3,0,40,257,110]
s: 327
d: [2,229,1,15,1,0,44,100,100]
s: 229
d: [2,345,3,14,5,0,44,254,327]
d: [2,352,3,12,4,0,18,345]
d: [2,363,3,9,3,0,21,352,110]
a: [3,366,3,2,1,363]
s: 366
r: 185
r: 175
r: 160
r: 246
r: 93
r: 113
r: 112
r: 278
r: 239
r: 327
r: 224
r: 215
r: 119
r: 110
r: 109
r: 98
r: 97
r: 95
d: [2,389,4,0,0,2,40,366,145]
e: 389
d: [2,388,4,0,0,2,40,366,165]
e: 388
d: [2,387,4,0,0,2,52,165,366]
e: 387
r: 366
s: 160
s: 175
s: 246
s: 185
a: [3,249,3,2,1,242]
s: 249
r: 273
r: 259
r: 257
r: 254
r: 252
r: 251
r: 117
r: 116
r: 106
r: 105
r: 104
r: 99
d: [2,398,4,0,0,2,40,249,180]
e: 398
d: [2,397,4,0,0,2,40,249,205]
e: 397
d: [2,396,4,0,0,1,40,249,114]
e: 396
d: [2,395,4,0,0,1,40,249,121]
e: 395
d: [2,394,4,0,0,2,40,249,210]
e: 394
d: [2,393,4,0,0,2,52,210,249]
e: 393
r: 249
s: 93
s: 95
s: 239
s: 97
s: 215
s: 112
s: 113
d: [2,411,1,4,1,0,44,112,100]
s: 411
s: 98
d: [2,412,2,4,1,0,44,112,227]
s: 412
s: 110
s: 116
d: [2,414,2,4,1,0,44,112,230]
s: 414
s: 109
d: [2,415,3,4,1,0,44,112,264]
s: 415
s: 119
s: 117
a: [3,369,3,5,1,363]
s: 369
s: 99
d: [2,416,1,7,3,0,44,98,98]
s: 416
s: 254
s: 259
d: [2,413,2,7,1,0,44,112,229]
s: 413
s: 273
d: [2,342,3,11,4,0,44,117,327]
d: [2,355,3,9,3,0,18,342]
a: [3,361,3,7,2,355]
s: 361
s: 104
s: 105
d: [2,474,4,0,0,3,52,210,210,135,369,369,105]
e: 474
f: 474,395,387,363,201,124,91,89
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t107_relat_1
% SZS output start Proof for relat_1__t107_relat_1
1. ! [X0] : k1_xboole_0 = k6_relat_1(X0,k1_xboole_0) [input t107_relat_1]
2. ~! [X0] : k1_xboole_0 = k6_relat_1(X0,k1_xboole_0) [negated conjecture 1]
4. ! [X0] : (v1_xboole_0(X0) => v1_relat_1(X0)) [input cc1_relat_1]
8. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (k6_relat_1(X0,X1) = X2 <=> ! [X3] : ! [X4] : (r2_hidden(k4_tarski(X3,X4),X2) <=> (r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X0)))))) [input d12_relat_1]
15. ! [X0,X1] : (v1_relat_1(X1) => v1_relat_1(k6_relat_1(X0,X1))) [input dt_k6_relat_1]
20. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
27. ! [X0,X1,X2,X3] : v1_relat_1(k2_tarski(k4_tarski(X0,X1),k4_tarski(X2,X3))) [input fc7_relat_1]
29. ? [X0] : v1_xboole_0(X0) [input rc1_xboole_0]
37. ! [X0] : (v1_relat_1(X0) => (! [X1] : ! [X2] : ~r2_hidden(k4_tarski(X1,X2),X0) => k1_xboole_0 = X0)) [input t37_relat_1]
39. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
41. ! [X0] : (v1_relat_1(X0) => (! [X1,X2] : ~r2_hidden(k4_tarski(X1,X2),X0) => k1_xboole_0 = X0)) [flattening 37]
43. ! [X0,X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (k6_relat_1(X0,X1) = X2 <=> ! [X3,X4] : (r2_hidden(k4_tarski(X3,X4),X2) <=> (r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X0)))))) [flattening 8]
47. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 39]
50. ? [X0] : k1_xboole_0 != k6_relat_1(X0,k1_xboole_0) [ennf transformation 2]
53. ! [X0] : (v1_relat_1(X0) | ~v1_xboole_0(X0)) [ennf transformation 4]
55. ! [X0] : ((k1_xboole_0 = X0 | ? [X1,X2] : r2_hidden(k4_tarski(X1,X2),X0)) | ~v1_relat_1(X0)) [ennf transformation 41]
56. ! [X0] : (k1_xboole_0 = X0 | ? [X1,X2] : r2_hidden(k4_tarski(X1,X2),X0) | ~v1_relat_1(X0)) [flattening 55]
57. ! [X0,X1] : (! [X2] : ((k6_relat_1(X0,X1) = X2 <=> ! [X3,X4] : (r2_hidden(k4_tarski(X3,X4),X2) <=> (r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X0)))) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [ennf transformation 43]
62. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 47]
63. ! [X0,X1] : (v1_relat_1(k6_relat_1(X0,X1)) | ~v1_relat_1(X1)) [ennf transformation 15]
68. ? [X0] : k1_xboole_0 != k6_relat_1(X0,k1_xboole_0) => k1_xboole_0 != k6_relat_1(sK0,k1_xboole_0) [choice axiom]
69. k1_xboole_0 != k6_relat_1(sK0,k1_xboole_0) [skolemisation 50,68]
70. ! [X0] : (? [X1,X2] : r2_hidden(k4_tarski(X1,X2),X0) => r2_hidden(k4_tarski(sK1(X0),sK2(X0)),X0)) [choice axiom]
71. ! [X0] : (k1_xboole_0 = X0 | r2_hidden(k4_tarski(sK1(X0),sK2(X0)),X0) | ~v1_relat_1(X0)) [skolemisation 56,70]
72. ! [X0,X1] : (! [X2] : (((k6_relat_1(X0,X1) = X2 | ? [X3,X4] : (((~r2_hidden(k4_tarski(X3,X4),X1) | ~r2_hidden(X4,X0)) | ~r2_hidden(k4_tarski(X3,X4),X2)) & ((r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X0)) | r2_hidden(k4_tarski(X3,X4),X2)))) & (! [X3,X4] : ((r2_hidden(k4_tarski(X3,X4),X2) | (~r2_hidden(k4_tarski(X3,X4),X1) | ~r2_hidden(X4,X0))) & ((r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X0)) | ~r2_hidden(k4_tarski(X3,X4),X2))) | k6_relat_1(X0,X1) != X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [nnf transformation 57]
73. ! [X0,X1] : (! [X2] : (((k6_relat_1(X0,X1) = X2 | ? [X3,X4] : ((~r2_hidden(k4_tarski(X3,X4),X1) | ~r2_hidden(X4,X0) | ~r2_hidden(k4_tarski(X3,X4),X2)) & ((r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X0)) | r2_hidden(k4_tarski(X3,X4),X2)))) & (! [X3,X4] : ((r2_hidden(k4_tarski(X3,X4),X2) | ~r2_hidden(k4_tarski(X3,X4),X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X0)) | ~r2_hidden(k4_tarski(X3,X4),X2))) | k6_relat_1(X0,X1) != X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [flattening 72]
74. ! [X0,X1] : (! [X2] : (((k6_relat_1(X0,X1) = X2 | ? [X3,X4] : ((~r2_hidden(k4_tarski(X3,X4),X1) | ~r2_hidden(X4,X0) | ~r2_hidden(k4_tarski(X3,X4),X2)) & ((r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X0)) | r2_hidden(k4_tarski(X3,X4),X2)))) & (! [X5,X6] : ((r2_hidden(k4_tarski(X5,X6),X2) | ~r2_hidden(k4_tarski(X5,X6),X1) | ~r2_hidden(X6,X0)) & ((r2_hidden(k4_tarski(X5,X6),X1) & r2_hidden(X6,X0)) | ~r2_hidden(k4_tarski(X5,X6),X2))) | k6_relat_1(X0,X1) != X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [rectify 73]
75. ! [X2,X1,X0] : (? [X3,X4] : ((~r2_hidden(k4_tarski(X3,X4),X1) | ~r2_hidden(X4,X0) | ~r2_hidden(k4_tarski(X3,X4),X2)) & ((r2_hidden(k4_tarski(X3,X4),X1) & r2_hidden(X4,X0)) | r2_hidden(k4_tarski(X3,X4),X2))) => ((~r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X1) | ~r2_hidden(sK4(X0,X1,X2),X0) | ~r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X2)) & ((r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X1) & r2_hidden(sK4(X0,X1,X2),X0)) | r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X2)))) [choice axiom]
76. ! [X0,X1] : (! [X2] : (((k6_relat_1(X0,X1) = X2 | ((~r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X1) | ~r2_hidden(sK4(X0,X1,X2),X0) | ~r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X2)) & ((r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X1) & r2_hidden(sK4(X0,X1,X2),X0)) | r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X2)))) & (! [X5,X6] : ((r2_hidden(k4_tarski(X5,X6),X2) | ~r2_hidden(k4_tarski(X5,X6),X1) | ~r2_hidden(X6,X0)) & ((r2_hidden(k4_tarski(X5,X6),X1) & r2_hidden(X6,X0)) | ~r2_hidden(k4_tarski(X5,X6),X2))) | k6_relat_1(X0,X1) != X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [skolemisation 74,75]
83. ? [X0] : v1_xboole_0(X0) => v1_xboole_0(sK8) [choice axiom]
84. v1_xboole_0(sK8) [skolemisation 29,83]
89. k1_xboole_0 != k6_relat_1(sK0,k1_xboole_0) [cnf transformation 69] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
91. v1_xboole_0(k1_xboole_0) [cnf transformation 20] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
97. ~v1_xboole_0(X0) | v1_relat_1(X0) [cnf transformation 53] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
99. r2_hidden(k4_tarski(sK1(X0),sK2(X0)),X0) | k1_xboole_0 = X0 | ~v1_relat_1(X0) [cnf transformation 71] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
101. r2_hidden(X6,X0) | ~r2_hidden(k4_tarski(X5,X6),X2) | k6_relat_1(X0,X1) != X2 | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 76] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
105. r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X2) | r2_hidden(k4_tarski(sK3(X0,X1,X2),sK4(X0,X1,X2)),X1) | k6_relat_1(X0,X1) = X2 | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 76] {a:0,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
110. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 62] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
116. v1_relat_1(k6_relat_1(X0,X1)) | ~v1_relat_1(X1) [cnf transformation 63] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
121. v1_relat_1(k2_tarski(k4_tarski(X0,X1),k4_tarski(X2,X3))) [cnf transformation 27] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
124. v1_xboole_0(sK8) [cnf transformation 84] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
131. r2_hidden(X6,X0) | ~r2_hidden(k4_tarski(X5,X6),k6_relat_1(X0,X1)) | ~v1_relat_1(k6_relat_1(X0,X1)) | ~v1_relat_1(X1) [equality resolution 101] {a:0,w:16,wCS:6400,sine:1,thAx:0,allAx:2,thDist:-2}
133. 1 <=> k1_xboole_0 = k6_relat_1(sK0,k1_xboole_0) [avatar definition]
135. k1_xboole_0 != k6_relat_1(sK0,k1_xboole_0) <- (~1) [avatar component clause 133] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
136. ~1 [avatar split clause 89,133]
143. 3 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
145. v1_xboole_0(k1_xboole_0) <- (3) [avatar component clause 143] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
146. 3 [avatar split clause 91,143]
163. 7 <=> v1_xboole_0(sK8) [avatar definition]
165. v1_xboole_0(sK8) <- (7) [avatar component clause 163] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
166. 7 [avatar split clause 124,163]
201. v1_relat_1(k1_xboole_0) <- (3) [resolution 97,145] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
208. 15 <=> v1_relat_1(k1_xboole_0) [avatar definition]
210. v1_relat_1(k1_xboole_0) <- (15) [avatar component clause 208] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
211. 15 | ~3 [avatar split clause 201,143,208]
248. 17 <=> ! [X0] : ~v1_relat_1(X0) [avatar definition]
249. ~v1_relat_1(X0) <- (17) [avatar component clause 248] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
254. ~r2_hidden(k4_tarski(X5,X6),k6_relat_1(X0,X1)) | r2_hidden(X6,X0) | ~v1_relat_1(X1) [subsumption resolution 131,116] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
255. r2_hidden(sK2(k6_relat_1(X0,X1)),X0) | ~v1_relat_1(X1) | k1_xboole_0 = k6_relat_1(X0,X1) | ~v1_relat_1(k6_relat_1(X0,X1)) [resolution 254,99] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:4,thDist:-4}
257. r2_hidden(sK2(k6_relat_1(X0,X1)),X0) | ~v1_relat_1(X1) | k1_xboole_0 = k6_relat_1(X0,X1) [subsumption resolution 255,116] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
327. k1_xboole_0 = k6_relat_1(X5,X4) | ~v1_relat_1(X4) | ~v1_xboole_0(X5) [resolution 257,110] {a:2,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
345. ~r2_hidden(k4_tarski(X12,X13),k1_xboole_0) | r2_hidden(X13,X10) | ~v1_relat_1(X11) | ~v1_relat_1(X11) | ~v1_xboole_0(X10) [superposition 254,327] {a:3,w:14,wCS:5600,sine:1,thAx:0,allAx:7,thDist:-7}
352. ~r2_hidden(k4_tarski(X12,X13),k1_xboole_0) | r2_hidden(X13,X10) | ~v1_relat_1(X11) | ~v1_xboole_0(X10) [duplicate literal removal 345] {a:3,w:12,wCS:4800,sine:1,thAx:0,allAx:7,thDist:-7}
363. ~r2_hidden(k4_tarski(X12,X13),k1_xboole_0) | ~v1_relat_1(X11) | ~v1_xboole_0(X10) [subsumption resolution 352,110] {a:3,w:9,wCS:3600,sine:1,thAx:0,allAx:7,thDist:-7}
365. 19 <=> ! [X10] : ~v1_xboole_0(X10) [avatar definition]
366. ~v1_xboole_0(X10) <- (19) [avatar component clause 365] {a:3,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:7,thDist:-7}
368. 20 <=> ! [X13,X12] : ~r2_hidden(k4_tarski(X12,X13),k1_xboole_0) [avatar definition]
369. ~r2_hidden(k4_tarski(X12,X13),k1_xboole_0) <- (20) [avatar component clause 368] {a:3,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:7,thDist:-7}
370. 19 | 17 | 20 [avatar split clause 363,368,248,365]
387. $false <- (7, 19) [unit resulting resolution 165,366] {a:4,w:0,sine:255,thAx:0,allAx:9,thDist:-9}
392. ~7 | ~19 [avatar contradiction clause 387]
395. $false <- (17) [resolution 249,121] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
402. ~17 [avatar contradiction clause 395]
474. $false <- (~1, 15, 20) [unit resulting resolution 210,210,135,369,369,105] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:22,thDist:-22}
483. 1 | ~15 | ~20 [avatar contradiction clause 474]
488. $false [avatar sat refutation 136,146,166,211,370,392,402,483] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:172,thDist:-172}
% SZS output end Proof for relat_1__t107_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 39
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 40
% Generated clauses: 338
% Active clauses: 84
% Passive clauses: 175
% Final active clauses: 51
% Final passive clauses: 74
% Main loop iterations started: 101
% 
% >>> Simplifying Inferences
% Duplicate literals: 28
% Trivial inequalities: 5
% Fw subsumption resolutions: 24
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 19
% Equational tautologies: 1
% Forward subsumptions: 73
% Backward subsumptions: 30
% 
% >>> Generating Inferences
% Binary resolution: 80
% Unit resulting resolution: 3
% Factoring: 2
% Forward superposition: 46
% Backward superposition: 55
% Self superposition: 9
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 9
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 31
% SAT solver unit clauses: 13
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 5117
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.004 s ( own 0.0000 s ) 

