i: [1,49,0,12,3,1,0,0] t56_zfmisc_1
a: [3,84,0,3,1,49]
s: 84
a: [3,88,0,3,1,49]
s: 88
a: [3,81,0,6,1,49]
s: 81
i: [1,50,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,96,0,2,1,50]
s: 96
i: [1,51,0,2,1,0,0,2] fc1_xboole_0
a: [3,101,0,2,1,51]
s: 101
i: [1,70,0,2,1,0,0,2] rc2_xboole_0
a: [3,111,0,2,1,70]
s: 111
i: [1,71,0,2,1,0,0,2] rc1_xboole_0
a: [3,116,0,2,1,71]
s: 116
i: [1,53,0,3,1,0,0,1] fc2_xboole_0
s: 53
i: [1,52,0,3,1,0,0,255] d2_xboole_0
a: [3,106,0,3,1,52]
s: 106
r: 96
i: [1,73,0,4,1,0,0,1] d1_tarski
s: 73
i: [1,54,0,5,1,0,0,1] t4_boole
s: 54
i: [1,55,0,5,1,0,0,1] t3_boole
s: 55
i: [1,56,0,5,2,0,0,2] t6_boole
s: 56
i: [1,62,0,5,2,0,0,1] t7_boole
s: 62
i: [1,69,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 69
i: [1,61,0,7,3,0,0,1] t8_boole
s: 61
d: [2,123,1,2,1,1,40,62,84]
a: [3,128,1,2,1,123]
s: 128
i: [1,74,0,7,2,0,0,1] d1_tarski
s: 74
d: [2,130,1,3,1,1,40,69,84]
a: [3,135,1,3,1,130]
s: 135
i: [1,76,0,8,2,0,0,1] d5_xboole_0
s: 76
i: [1,77,0,8,2,0,0,1] d5_xboole_0
s: 77
d: [2,131,1,4,1,0,40,69,73]
s: 131
i: [1,75,0,11,3,0,0,1] d5_xboole_0
s: 75
d: [2,147,1,7,2,1,40,75,81]
d: [2,150,1,4,1,2,21,147,84]
a: [3,154,1,4,1,150]
s: 154
d: [2,156,2,0,0,2,52,88,154,74]
e: 156
d: [2,157,2,3,1,1,40,154,74]
d: [2,161,2,0,0,2,21,157,88]
e: 161
r: 81
r: 88
a: [3,89,0,3,1,49]
s: 89
r: 154
a: [3,80,0,6,1,49]
s: 80
r: 84
r: 89
a: [3,85,0,3,1,49]
s: 85
s: 88
a: [3,153,1,4,1,150]
s: 153
s: 80
d: [2,176,1,0,0,2,52,85,80,77]
e: 176
d: [2,177,1,3,1,1,40,80,77]
d: [2,182,1,0,0,2,21,177,85]
e: 182
r: 85
r: 88
s: 84
s: 89
r: 80
r: 153
d: [2,198,1,4,1,2,23,153,89]
d: [2,200,1,0,0,2,21,198,73]
e: 200
d: [2,178,1,4,1,1,40,80,76]
d: [2,202,1,4,1,2,22,178,89]
d: [2,203,1,0,0,2,21,202,73]
e: 203
i: [1,48,0,9,2,1,0,0] t56_zfmisc_1
i: [1,47,0,9,2,1,0,0] t56_zfmisc_1
f: 203,177,157,150,47,48,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t56_zfmisc_1
% SZS output start Proof for zfmisc_1__t56_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2))) <=> (X0 != X2 & r2_hidden(X0,X1))) [input t56_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2))) <=> (X0 != X2 & r2_hidden(X0,X1))) [negated conjecture 1]
4. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
6. ! [X0] : ! [X1] : ! [X2] : (k4_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (~r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [input d5_xboole_0]
20. ~! [X0,X1,X2] : (r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2))) <=> (X0 != X2 & r2_hidden(X0,X1))) [flattening 2]
21. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 4]
24. ! [X0,X1,X2] : (k4_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (~r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [flattening 6]
25. ? [X0,X1,X2] : (r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2))) <~> (X0 != X2 & r2_hidden(X0,X1))) [ennf transformation 20]
30. ? [X0,X1,X2] : (((X0 = X2 | ~r2_hidden(X0,X1)) | ~r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2)))) & ((X0 != X2 & r2_hidden(X0,X1)) | r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2))))) [nnf transformation 25]
31. ? [X0,X1,X2] : ((X0 = X2 | ~r2_hidden(X0,X1) | ~r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2)))) & ((X0 != X2 & r2_hidden(X0,X1)) | r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2))))) [flattening 30]
32. ? [X0,X1,X2] : ((X0 = X2 | ~r2_hidden(X0,X1) | ~r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2)))) & ((X0 != X2 & r2_hidden(X0,X1)) | r2_hidden(X0,k4_xboole_0(X1,k1_tarski(X2))))) => ((sK0 = sK2 | ~r2_hidden(sK0,sK1) | ~r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2)))) & ((sK0 != sK2 & r2_hidden(sK0,sK1)) | r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2))))) [choice axiom]
33. (sK0 = sK2 | ~r2_hidden(sK0,sK1) | ~r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2)))) & ((sK0 != sK2 & r2_hidden(sK0,sK1)) | r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2)))) [skolemisation 31,32]
34. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 21]
35. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 34]
36. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
37. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 35,36]
38. ! [X0,X1,X2] : ((k4_xboole_0(X0,X1) = X2 | ? [X3] : (((r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (r2_hidden(X3,X1) | ~r2_hidden(X3,X0))) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k4_xboole_0(X0,X1) != X2)) [nnf transformation 24]
39. ! [X0,X1,X2] : ((k4_xboole_0(X0,X1) = X2 | ? [X3] : ((r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k4_xboole_0(X0,X1) != X2)) [flattening 38]
40. ! [X0,X1,X2] : ((k4_xboole_0(X0,X1) = X2 | ? [X3] : ((r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((~r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k4_xboole_0(X0,X1) != X2)) [rectify 39]
41. ! [X2,X1,X0] : (? [X3] : ((r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((~r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2))) => ((r2_hidden(sK4(X0,X1,X2),X1) | ~r2_hidden(sK4(X0,X1,X2),X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & ((~r2_hidden(sK4(X0,X1,X2),X1) & r2_hidden(sK4(X0,X1,X2),X0)) | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
42. ! [X0,X1,X2] : ((k4_xboole_0(X0,X1) = X2 | ((r2_hidden(sK4(X0,X1,X2),X1) | ~r2_hidden(sK4(X0,X1,X2),X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & ((~r2_hidden(sK4(X0,X1,X2),X1) & r2_hidden(sK4(X0,X1,X2),X0)) | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((~r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k4_xboole_0(X0,X1) != X2)) [skolemisation 40,41]
47. r2_hidden(sK0,sK1) | r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2))) [cnf transformation 33] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. sK0 != sK2 | r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2))) [cnf transformation 33] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. sK0 = sK2 | ~r2_hidden(sK0,sK1) | ~r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2))) [cnf transformation 33] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
57. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 37] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
58. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 37] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
63. r2_hidden(X4,X0) | ~r2_hidden(X4,X2) | k4_xboole_0(X0,X1) != X2 [cnf transformation 42] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
64. ~r2_hidden(X4,X1) | ~r2_hidden(X4,X2) | k4_xboole_0(X0,X1) != X2 [cnf transformation 42] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
65. r2_hidden(X4,X2) | r2_hidden(X4,X1) | ~r2_hidden(X4,X0) | k4_xboole_0(X0,X1) != X2 [cnf transformation 42] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
72. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 58] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:2,thDist:-2}
73. r2_hidden(X3,k1_tarski(X3)) [equality resolution 72] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
74. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 57] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
75. r2_hidden(X4,k4_xboole_0(X0,X1)) | r2_hidden(X4,X1) | ~r2_hidden(X4,X0) [equality resolution 65] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
76. ~r2_hidden(X4,k4_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) [equality resolution 64] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
77. ~r2_hidden(X4,k4_xboole_0(X0,X1)) | r2_hidden(X4,X0) [equality resolution 63] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
79. 1 <=> r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2))) [avatar definition]
80. r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2))) <- (1) [avatar component clause 79] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
81. ~r2_hidden(sK0,k4_xboole_0(sK1,k1_tarski(sK2))) <- (~1) [avatar component clause 79] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. 2 <=> r2_hidden(sK0,sK1) [avatar definition]
84. r2_hidden(sK0,sK1) <- (2) [avatar component clause 83] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. 3 <=> sK0 = sK2 [avatar definition]
89. sK0 = sK2 <- (3) [avatar component clause 87] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
90. ~1 | ~2 | 3 [avatar split clause 49,87,83,79]
91. 1 | ~3 [avatar split clause 48,87,79]
92. 1 | 2 [avatar split clause 47,83,79]
147. r2_hidden(sK0,k1_tarski(sK2)) | ~r2_hidden(sK0,sK1) <- (~1) [resolution 75,81] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
150. r2_hidden(sK0,k1_tarski(sK2)) <- (~1, 2) [subsumption resolution 147,84] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
152. 11 <=> r2_hidden(sK0,k1_tarski(sK2)) [avatar definition]
154. r2_hidden(sK0,k1_tarski(sK2)) <- (11) [avatar component clause 152] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
155. 11 | 1 | ~2 [avatar split clause 150,83,79,152]
157. sK0 = sK2 <- (11) [resolution 154,74] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
163. 3 | ~11 [avatar split clause 157,152,87]
177. r2_hidden(sK0,sK1) <- (1) [resolution 80,77] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
178. ~r2_hidden(sK0,k1_tarski(sK2)) <- (1) [resolution 80,76] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
184. 2 | ~1 [avatar split clause 177,79,83]
202. ~r2_hidden(sK0,k1_tarski(sK0)) <- (1, 3) [forward demodulation 178,89] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
203. $false <- (1, 3) [subsumption resolution 202,73] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
204. ~1 | ~3 [avatar contradiction clause 203]
205. $false [avatar sat refutation 90,91,92,155,163,184,204] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:115,thDist:-115}
% SZS output end Proof for zfmisc_1__t56_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 25
% Generated clauses: 97
% Active clauses: 31
% Passive clauses: 49
% Final active clauses: 21
% Final passive clauses: 17
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 5
% Fw demodulations: 2
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Equational tautologies: 1
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 2
% Forward superposition: 6
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 5
% SAT solver binary clauses: 15
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward demodulation: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

