i: [1,183,0,2,1,1,0,0] t11_xxreal_1
a: [3,300,0,2,1,183]
s: 300
i: [1,182,0,2,1,1,0,0] t11_xxreal_1
a: [3,305,0,2,1,182]
s: 305
i: [1,186,0,3,1,1,0,0] t11_xxreal_1
a: [3,281,0,3,1,186]
s: 281
i: [1,185,0,8,2,1,0,0] t11_xxreal_1
a: [3,286,0,5,1,185]
s: 286
i: [1,184,0,5,1,1,0,0] t11_xxreal_1
a: [3,295,0,5,1,184]
s: 295
i: [1,187,0,2,1,0,0,3] fc5_numbers
a: [3,310,0,2,1,187]
s: 310
i: [1,188,0,2,1,0,0,2] fc2_membered
a: [3,315,0,2,1,188]
s: 315
i: [1,189,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,320,0,2,1,189]
s: 320
i: [1,190,0,2,1,0,0,255] fc1_xboole_0
a: [3,325,0,2,1,190]
s: 325
i: [1,237,0,7,3,0,0,2] reflexivity_r1_xxreal_0
a: [3,334,0,2,1,237]
s: 334
d: [2,466,1,0,0,2,40,334,305]
e: 466
d: [2,465,1,0,0,2,40,334,300]
e: 465
d: [2,464,1,0,0,2,52,300,334]
e: 464
r: 334
i: [1,251,0,2,1,0,0,2] rc2_xboole_0
a: [3,342,0,2,1,251]
s: 342
i: [1,252,0,2,1,0,0,255] rc1_xreal_0
a: [3,347,0,2,1,252]
s: 347
i: [1,253,0,2,1,0,0,2] rc1_xboole_0
a: [3,352,0,2,1,253]
s: 352
i: [1,254,0,2,1,0,0,1] rc1_xxreal_0
a: [3,357,0,2,1,254]
s: 357
i: [1,256,0,2,1,0,0,255] rc1_membered
a: [3,362,0,2,1,256]
s: 362
i: [1,255,0,2,1,0,0,255] rc1_membered
a: [3,367,0,2,1,255]
s: 367
i: [1,258,0,2,1,0,0,255] rc2_membered
a: [3,372,0,2,1,258]
s: 372
i: [1,257,0,2,1,0,0,255] rc2_membered
a: [3,377,0,2,1,257]
s: 377
i: [1,260,0,2,1,0,0,255] rc3_membered
a: [3,382,0,2,1,260]
s: 382
i: [1,259,0,2,1,0,0,255] rc3_membered
a: [3,387,0,2,1,259]
s: 387
i: [1,263,0,2,1,0,0,255] rc3_xreal_0
a: [3,392,0,2,1,263]
s: 392
i: [1,262,0,2,1,0,0,255] rc3_xreal_0
a: [3,397,0,2,1,262]
s: 397
i: [1,261,0,2,1,0,0,255] rc3_xreal_0
a: [3,402,0,2,1,261]
s: 402
i: [1,266,0,2,1,0,0,255] rc2_xreal_0
a: [3,407,0,2,1,266]
s: 407
i: [1,265,0,2,1,0,0,255] rc2_xreal_0
a: [3,412,0,2,1,265]
s: 412
i: [1,264,0,2,1,0,0,255] rc2_xreal_0
a: [3,417,0,2,1,264]
s: 417
i: [1,269,0,2,1,0,0,255] rc4_xreal_0
a: [3,422,0,2,1,269]
s: 422
i: [1,268,0,2,1,0,0,255] rc4_xreal_0
a: [3,427,0,2,1,268]
s: 427
i: [1,267,0,2,1,0,0,255] rc4_xreal_0
a: [3,432,0,2,1,267]
s: 432
i: [1,271,0,2,1,0,0,2] rc4_xxreal_0
a: [3,437,0,2,1,271]
s: 437
i: [1,270,0,2,1,0,0,2] rc4_xxreal_0
a: [3,442,0,2,1,270]
s: 442
i: [1,273,0,2,1,0,0,255] rc3_xxreal_0
a: [3,447,0,2,1,273]
s: 447
i: [1,272,0,2,1,0,0,255] rc3_xxreal_0
a: [3,452,0,2,1,272]
s: 452
i: [1,275,0,2,1,0,0,255] rc2_xxreal_0
a: [3,457,0,2,1,275]
s: 457
i: [1,274,0,2,1,0,0,255] rc2_xxreal_0
a: [3,462,0,2,1,274]
s: 462
i: [1,191,0,3,1,0,0,255] d2_xboole_0
a: [3,330,0,3,1,191]
s: 330
r: 320
i: [1,192,0,4,2,0,0,255] cc2_membered
s: 192
i: [1,193,0,4,2,0,0,255] cc3_membered
s: 193
i: [1,194,0,4,2,0,0,255] cc2_xreal_0
s: 194
i: [1,195,0,4,2,0,0,255] cc2_xxreal_0
s: 195
i: [1,196,0,4,2,0,0,255] cc4_membered
s: 196
i: [1,199,0,4,2,0,0,255] cc1_membered
s: 199
i: [1,201,0,4,2,0,0,255] cc4_xreal_0
s: 201
d: [2,471,1,2,1,1,40,199,362]
a: [3,477,1,2,1,471]
s: 477
i: [1,202,0,4,2,0,0,255] cc18_membered
s: 202
i: [1,235,0,4,1,0,0,2] existence_m1_subset_1
s: 235
d: [2,472,1,2,1,1,40,199,372]
a: [3,482,1,2,1,472]
s: 482
i: [1,203,0,5,2,0,0,255] t6_boole
s: 203
d: [2,473,1,2,1,1,40,199,382]
a: [3,487,1,2,1,473]
s: 487
i: [1,213,0,5,2,0,0,3] cc1_xxreal_0
s: 213
i: [1,232,0,5,2,0,0,1] t7_boole
s: 232
d: [2,548,1,4,1,1,40,232,295]
a: [3,552,1,4,1,548]
s: 552
a: [3,337,0,5,2,237]
s: 337
d: [2,489,1,2,1,1,40,201,347]
a: [3,496,1,2,1,489]
s: 496
i: [1,219,0,6,3,0,0,255] cc7_xxreal_0
s: 219
i: [1,220,0,6,3,0,0,255] cc7_xxreal_0
s: 220
d: [2,504,1,2,1,1,40,202,325]
a: [3,511,1,2,1,504]
s: 511
i: [1,223,0,6,3,0,0,255] cc5_xxreal_0
s: 223
d: [2,505,1,2,1,1,40,202,352]
a: [3,516,1,2,1,505]
s: 516
i: [1,225,0,6,3,0,0,255] cc5_xxreal_0
s: 225
i: [1,226,0,6,3,0,0,255] cc3_xxreal_0
s: 226
d: [2,506,1,2,1,1,40,202,432]
a: [3,521,1,2,1,506]
s: 521
i: [1,228,0,6,3,0,0,255] cc3_xxreal_0
s: 228
d: [2,507,1,2,1,1,40,202,442]
a: [3,526,1,2,1,507]
s: 526
i: [1,230,0,6,2,0,0,1] t1_subset
s: 230
i: [1,236,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 236
d: [2,621,1,5,1,1,40,230,295]
a: [3,625,1,5,1,621]
s: 625
i: [1,197,0,7,3,0,0,255] cc14_membered
s: 197
d: [2,634,2,6,2,1,40,197,625]
a: [3,638,2,4,1,634]
s: 638
i: [1,198,0,7,3,0,0,2] cc13_membered
s: 198
i: [1,200,0,7,3,0,0,255] cc17_membered
s: 200
d: [2,645,2,6,2,1,40,198,625]
a: [3,649,2,4,1,645]
s: 649
i: [1,231,0,7,3,0,0,1] t8_boole
s: 231
d: [2,656,2,6,2,1,40,200,625]
a: [3,660,2,4,1,656]
s: 660
i: [1,233,0,7,2,0,0,1] t2_tarski
s: 233
i: [1,215,0,8,4,0,0,255] cc4_xxreal_0
s: 215
d: [2,669,1,6,3,1,40,215,305]
a: [3,682,1,2,1,669]
s: 682
r: 682
i: [1,217,0,8,4,0,0,255] cc6_xxreal_0
s: 217
a: [3,681,1,2,1,669]
s: 681
i: [1,221,0,8,4,0,0,255] cc8_xxreal_0
s: 221
i: [1,229,0,8,3,0,0,1] t2_subset
s: 229
a: [3,686,1,2,1,669]
s: 686
i: [1,238,0,8,3,0,0,1] fc2_xxreal_1
s: 238
d: [2,670,1,6,3,1,40,215,300]
a: [3,695,1,2,1,670]
s: 695
r: 695
i: [1,239,0,8,3,0,0,1] fc3_xxreal_1
s: 239
d: [2,746,3,0,0,3,52,305,300,649,239]
e: 746
d: [2,747,3,4,2,1,40,239,649]
d: [2,749,3,2,1,2,21,747,300]
d: [2,750,3,0,0,3,21,749,305]
e: 750
r: 649
i: [1,234,0,10,3,0,0,1] t2_tarski
s: 234
a: [3,694,1,2,1,670]
s: 694
i: [1,240,0,10,4,0,0,2] connectedness_r1_xxreal_0
s: 240
a: [3,699,1,2,1,670]
s: 699
i: [1,204,0,11,3,0,0,255] d3_xxreal_1
s: 204
i: [1,205,0,11,3,0,0,255] d2_xxreal_1
s: 205
a: [3,653,2,2,1,645]
s: 653
i: [1,206,0,13,5,0,0,2] t1_xxreal_0
s: 206
d: [2,756,3,6,3,1,40,653,215]
a: [3,760,3,2,1,756]
s: 760
r: 760
i: [1,207,0,14,5,0,0,1] t2_xxreal_1
s: 207
i: [1,208,0,14,5,0,0,1] t2_xxreal_1
s: 208
a: [3,759,3,2,1,756]
s: 759
i: [1,210,0,14,5,0,0,1] t3_xxreal_1
s: 210
a: [3,764,3,2,1,756]
s: 764
i: [1,211,0,14,5,0,0,1] t3_xxreal_1
s: 211
i: [1,241,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 241
d: [2,780,1,9,4,1,40,207,295]
d: [2,782,1,7,3,2,21,780,300]
d: [2,783,1,5,2,3,21,782,305]
d: [2,784,1,3,1,4,21,783,653]
a: [3,788,1,3,1,784]
s: 788
i: [1,242,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 242
d: [2,790,1,9,4,1,40,208,295]
d: [2,792,1,7,3,2,21,790,300]
d: [2,793,1,5,2,3,21,792,305]
d: [2,794,1,3,1,4,21,793,653]
a: [3,798,1,3,1,794]
s: 798
i: [1,243,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 243
i: [1,244,0,15,4,0,0,255] fraenkel_a_2_2_xxreal_1
s: 244
d: [2,826,2,7,3,1,40,788,240]
d: [2,827,2,5,2,2,21,826,653]
d: [2,828,2,3,1,3,21,827,305]
a: [3,832,2,3,1,828]
s: 832
i: [1,246,0,15,4,0,0,255] fraenkel_a_2_1_xxreal_1
s: 246
d: [2,836,2,10,4,1,40,798,206]
d: [2,837,2,7,3,2,21,836,281]
d: [2,838,2,5,2,3,21,837,653]
d: [2,839,2,3,1,4,21,838,300]
a: [3,843,2,3,1,839]
s: 843
i: [1,247,0,15,4,0,0,255] fraenkel_a_2_1_xxreal_1
s: 247
i: [1,248,0,15,4,0,0,255] fraenkel_a_2_1_xxreal_1
s: 248
a: [3,648,2,4,1,645]
s: 648
i: [1,249,0,15,4,0,0,255] fraenkel_a_2_1_xxreal_1
s: 249
d: [2,627,1,5,1,1,40,236,295]
a: [3,631,1,5,1,627]
s: 631
i: [1,209,0,17,6,0,0,1] t2_xxreal_1
s: 209
i: [1,212,0,17,6,0,0,1] t3_xxreal_1
s: 212
d: [2,891,3,0,0,6,52,653,305,300,832,843,286,212]
e: 891
d: [2,892,1,12,5,1,40,212,286]
d: [2,901,1,9,4,2,21,892,843]
d: [2,902,1,6,3,3,21,901,832]
d: [2,903,1,4,2,4,21,902,300]
d: [2,904,1,2,1,5,21,903,305]
d: [2,905,1,0,0,6,21,904,653]
e: 905
f: 905,839,828,794,784,750,645,621,182,183,184,185,186
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_1__t11_xxreal_1
% SZS output start Proof for xxreal_1__t11_xxreal_1
1. ! [X0] : ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ~(X0 != X2 & ~(X0 != X1 & r2_hidden(X0,k2_xxreal_1(X1,X2))) & r2_hidden(X0,k3_xxreal_1(X1,X2))))) [input t11_xxreal_1]
2. ~! [X0] : ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ~(X0 != X2 & ~(X0 != X1 & r2_hidden(X0,k2_xxreal_1(X1,X2))) & r2_hidden(X0,k3_xxreal_1(X1,X2))))) [negated conjecture 1]
5. ! [X0] : (v2_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v1_xxreal_0(X1))) [input cc13_membered]
28. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
42. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v2_membered(k3_xxreal_1(X0,X1))) [input fc3_xxreal_1]
60. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
61. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input t1_xxreal_0]
64. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k3_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1)))))) [input t2_xxreal_1]
65. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r2_hidden(X0,k2_xxreal_1(X1,X2)) <=> (~r1_xxreal_0(X2,X0) & r1_xxreal_0(X1,X0)))))) [input t3_xxreal_1]
69. ~! [X0,X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ~(X0 != X2 & ~(X0 != X1 & r2_hidden(X0,k2_xxreal_1(X1,X2))) & r2_hidden(X0,k3_xxreal_1(X1,X2))))) [flattening 2]
71. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 60]
85. ? [X0,X1] : (? [X2] : ((X0 != X2 & (X0 = X1 | ~r2_hidden(X0,k2_xxreal_1(X1,X2))) & r2_hidden(X0,k3_xxreal_1(X1,X2))) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) [ennf transformation 69]
86. ? [X0,X1] : (? [X2] : (X0 != X2 & (X0 = X1 | ~r2_hidden(X0,k2_xxreal_1(X1,X2))) & r2_hidden(X0,k3_xxreal_1(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) [flattening 85]
93. ! [X0] : (! [X1] : (v1_xxreal_0(X1) | ~m1_subset_1(X1,X0)) | ~v2_membered(X0)) [ennf transformation 5]
101. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 61]
102. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 101]
103. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X0,k3_xxreal_1(X1,X2)) <=> (r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 64]
104. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X0,k2_xxreal_1(X1,X2)) <=> (~r1_xxreal_0(X2,X0) & r1_xxreal_0(X1,X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 65]
120. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 71]
129. ! [X0,X1] : (v2_membered(k3_xxreal_1(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 42]
130. ! [X0,X1] : (v2_membered(k3_xxreal_1(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 129]
131. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 28]
132. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 131]
137. ? [X0,X1] : (? [X2] : (X0 != X2 & (X0 = X1 | ~r2_hidden(X0,k2_xxreal_1(X1,X2))) & r2_hidden(X0,k3_xxreal_1(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (sK0 != X2 & (sK0 = sK1 | ~r2_hidden(sK0,k2_xxreal_1(sK1,X2))) & r2_hidden(sK0,k3_xxreal_1(sK1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
138. ? [X2] : (sK0 != X2 & (sK0 = sK1 | ~r2_hidden(sK0,k2_xxreal_1(sK1,X2))) & r2_hidden(sK0,k3_xxreal_1(sK1,X2)) & v1_xxreal_0(X2)) => (sK0 != sK2 & (sK0 = sK1 | ~r2_hidden(sK0,k2_xxreal_1(sK1,sK2))) & r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) & v1_xxreal_0(sK2)) [choice axiom]
139. (sK0 != sK2 & (sK0 = sK1 | ~r2_hidden(sK0,k2_xxreal_1(sK1,sK2))) & r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1) [skolemisation 86,138,137]
140. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k3_xxreal_1(X1,X2)) | (~r1_xxreal_0(X0,X2) | r1_xxreal_0(X0,X1))) & ((r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1)) | ~r2_hidden(X0,k3_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 103]
141. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k3_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X2) | r1_xxreal_0(X0,X1)) & ((r1_xxreal_0(X0,X2) & ~r1_xxreal_0(X0,X1)) | ~r2_hidden(X0,k3_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 140]
142. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k2_xxreal_1(X1,X2)) | (r1_xxreal_0(X2,X0) | ~r1_xxreal_0(X1,X0))) & ((~r1_xxreal_0(X2,X0) & r1_xxreal_0(X1,X0)) | ~r2_hidden(X0,k2_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 104]
143. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X0,k2_xxreal_1(X1,X2)) | r1_xxreal_0(X2,X0) | ~r1_xxreal_0(X1,X0)) & ((~r1_xxreal_0(X2,X0) & r1_xxreal_0(X1,X0)) | ~r2_hidden(X0,k2_xxreal_1(X1,X2)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 142]
182. v1_xxreal_0(sK1) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
183. v1_xxreal_0(sK2) [cnf transformation 139] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
184. r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) [cnf transformation 139] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
185. sK0 = sK1 | ~r2_hidden(sK0,k2_xxreal_1(sK1,sK2)) [cnf transformation 139] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
186. sK0 != sK2 [cnf transformation 139] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
198. ~m1_subset_1(X1,X0) | v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 93] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
206. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 102] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
207. ~r2_hidden(X0,k3_xxreal_1(X1,X2)) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 141] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
208. ~r2_hidden(X0,k3_xxreal_1(X1,X2)) | r1_xxreal_0(X0,X2) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 141] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
212. r2_hidden(X0,k2_xxreal_1(X1,X2)) | r1_xxreal_0(X2,X0) | ~r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 143] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
230. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 120] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
239. v2_membered(k3_xxreal_1(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 130] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
240. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 132] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
279. 1 <=> sK0 = sK2 [avatar definition]
281. sK0 != sK2 <- (~1) [avatar component clause 279] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
282. ~1 [avatar split clause 186,279]
284. 2 <=> r2_hidden(sK0,k2_xxreal_1(sK1,sK2)) [avatar definition]
286. ~r2_hidden(sK0,k2_xxreal_1(sK1,sK2)) <- (~2) [avatar component clause 284] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
288. 3 <=> sK0 = sK1 [avatar definition]
291. ~2 | 3 [avatar split clause 185,288,284]
293. 4 <=> r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) [avatar definition]
295. r2_hidden(sK0,k3_xxreal_1(sK1,sK2)) <- (4) [avatar component clause 293] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
296. 4 [avatar split clause 184,293]
298. 5 <=> v1_xxreal_0(sK2) [avatar definition]
300. v1_xxreal_0(sK2) <- (5) [avatar component clause 298] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
301. 5 [avatar split clause 183,298]
303. 6 <=> v1_xxreal_0(sK1) [avatar definition]
305. v1_xxreal_0(sK1) <- (6) [avatar component clause 303] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
306. 6 [avatar split clause 182,303]
621. m1_subset_1(sK0,k3_xxreal_1(sK1,sK2)) <- (4) [resolution 230,295] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
623. 62 <=> m1_subset_1(sK0,k3_xxreal_1(sK1,sK2)) [avatar definition]
625. m1_subset_1(sK0,k3_xxreal_1(sK1,sK2)) <- (62) [avatar component clause 623] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
626. 62 | ~4 [avatar split clause 621,293,623]
645. v1_xxreal_0(sK0) | ~v2_membered(k3_xxreal_1(sK1,sK2)) <- (62) [resolution 198,625] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
647. 66 <=> v2_membered(k3_xxreal_1(sK1,sK2)) [avatar definition]
649. ~v2_membered(k3_xxreal_1(sK1,sK2)) <- (~66) [avatar component clause 647] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
651. 67 <=> v1_xxreal_0(sK0) [avatar definition]
653. v1_xxreal_0(sK0) <- (67) [avatar component clause 651] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
654. ~66 | 67 | ~62 [avatar split clause 645,623,651,647]
747. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~66) [resolution 239,649] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
749. ~v1_xxreal_0(sK1) <- (5, ~66) [subsumption resolution 747,300] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
750. $false <- (5, 6, ~66) [subsumption resolution 749,305] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
751. ~5 | ~6 | 66 [avatar contradiction clause 750]
780. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4) [resolution 207,295] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
782. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4, 5) [subsumption resolution 780,300] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
783. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK0) <- (4, 5, 6) [subsumption resolution 782,305] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
784. ~r1_xxreal_0(sK0,sK1) <- (4, 5, 6, 67) [subsumption resolution 783,653] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
786. 88 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
788. ~r1_xxreal_0(sK0,sK1) <- (~88) [avatar component clause 786] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
789. ~88 | ~4 | ~5 | ~6 | ~67 [avatar split clause 784,651,303,298,293,786]
790. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4) [resolution 208,295] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
792. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4, 5) [subsumption resolution 790,300] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
793. r1_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK0) <- (4, 5, 6) [subsumption resolution 792,305] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
794. r1_xxreal_0(sK0,sK2) <- (4, 5, 6, 67) [subsumption resolution 793,653] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
796. 89 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
798. r1_xxreal_0(sK0,sK2) <- (89) [avatar component clause 796] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
799. 89 | ~4 | ~5 | ~6 | ~67 [avatar split clause 794,651,303,298,293,796]
826. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~88) [resolution 788,240] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
827. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) <- (67, ~88) [subsumption resolution 826,653] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
828. r1_xxreal_0(sK1,sK0) <- (6, 67, ~88) [subsumption resolution 827,305] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
830. 94 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
832. r1_xxreal_0(sK1,sK0) <- (94) [avatar component clause 830] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
833. 94 | ~6 | ~67 | 88 [avatar split clause 828,786,651,303,830]
836. sK0 = sK2 | ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) <- (89) [resolution 798,206] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
837. ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) <- (~1, 89) [subsumption resolution 836,281] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
838. ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) <- (~1, 67, 89) [subsumption resolution 837,653] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
839. ~r1_xxreal_0(sK2,sK0) <- (~1, 5, 67, 89) [subsumption resolution 838,300] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
841. 95 <=> r1_xxreal_0(sK2,sK0) [avatar definition]
843. ~r1_xxreal_0(sK2,sK0) <- (~95) [avatar component clause 841] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
844. ~95 | 1 | ~5 | ~67 | ~89 [avatar split clause 839,796,651,298,279,841]
892. r1_xxreal_0(sK2,sK0) | ~r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2) [resolution 212,286] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
901. ~r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, ~95) [subsumption resolution 892,843] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
902. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, 94, ~95) [subsumption resolution 901,832] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
903. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, 5, 94, ~95) [subsumption resolution 902,300] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
904. ~v1_xxreal_0(sK0) <- (~2, 5, 6, 94, ~95) [subsumption resolution 903,305] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
905. $false <- (~2, 5, 6, 67, 94, ~95) [subsumption resolution 904,653] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
906. 2 | ~5 | ~6 | ~67 | ~94 | 95 [avatar contradiction clause 905]
907. sK0 != sK1 | r1_xxreal_0(sK0,sK1) | ~r1_xxreal_0(sK1,sK0) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
908. $false [avatar sat refutation 282,291,296,301,306,626,654,751,789,799,833,844,906,907] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:452,thDist:-452}
% SZS output end Proof for xxreal_1__t11_xxreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 67
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 5
% 
% >>> Saturation
% Initial clauses: 94
% Generated clauses: 338
% Active clauses: 111
% Passive clauses: 169
% Final active clauses: 105
% Final passive clauses: 53
% Main loop iterations started: 116
% 
% >>> Simplifying Inferences
% Duplicate literals: 47
% Fw subsumption resolutions: 28
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 12
% Equational tautologies: 1
% Forward subsumptions: 24
% 
% >>> Generating Inferences
% Binary resolution: 99
% Unit resulting resolution: 3
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 14
% Split components: 33
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 90
% SAT solver unit clauses: 35
% SAT solver binary clauses: 30
% 
% Memory used [KB]: 5500
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

