i: [1,60,0,2,1,1,0,0] t122_funct_4
a: [3,105,0,2,1,60]
s: 105
i: [1,59,0,2,1,1,0,0] t122_funct_4
a: [3,110,0,2,1,59]
s: 110
i: [1,58,0,2,1,1,0,0] t122_funct_4
a: [3,115,0,2,1,58]
s: 115
i: [1,57,0,2,1,1,0,0] t122_funct_4
a: [3,120,0,2,1,57]
s: 120
i: [1,56,0,2,1,1,0,0] t122_funct_4
a: [3,125,0,2,1,56]
s: 125
i: [1,55,0,2,1,1,0,0] t122_funct_4
a: [3,130,0,2,1,55]
s: 130
i: [1,54,0,2,1,1,0,0] t122_funct_4
a: [3,135,0,2,1,54]
s: 135
i: [1,53,0,2,1,1,0,0] t122_funct_4
a: [3,140,0,2,1,53]
s: 140
i: [1,62,0,3,1,1,0,0] t122_funct_4
a: [3,95,0,3,1,62]
s: 95
i: [1,61,0,3,1,1,0,0] t122_funct_4
a: [3,100,0,3,1,61]
s: 100
i: [1,63,0,5,1,1,0,0] t122_funct_4
a: [3,90,0,5,1,63]
s: 90
i: [1,64,0,7,1,1,0,0] t122_funct_4
a: [3,85,0,7,1,64]
s: 85
i: [1,81,0,2,1,0,0,1] rc1_funct_1
a: [3,152,0,2,1,81]
s: 152
i: [1,80,0,2,1,0,0,1] rc1_funct_1
a: [3,157,0,2,1,80]
s: 157
i: [1,74,0,3,1,0,0,1] reflexivity_r1_tarski
s: 74
i: [1,73,0,4,1,0,0,255] existence_m1_subset_1
s: 73
i: [1,75,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 75
i: [1,70,0,7,2,0,0,255] t3_subset
s: 70
d: [2,159,1,5,1,1,40,75,90]
a: [3,163,1,5,1,159]
s: 163
i: [1,71,0,7,2,0,0,255] t3_subset
s: 71
i: [1,65,0,8,3,0,0,255] cc2_relat_1
s: 65
d: [2,165,1,5,1,0,40,70,73]
s: 165
i: [1,72,0,9,3,0,0,1] t1_xboole_1
s: 72
d: [2,171,1,6,2,1,40,72,100]
s: 171
i: [1,76,0,13,5,0,0,255] idempotence_k1_funct_4
a: [3,147,0,9,3,76]
s: 147
i: [1,66,0,10,4,0,0,255] cc3_funct_1
s: 66
d: [2,175,2,5,1,1,40,171,165]
a: [3,179,2,5,1,175]
s: 179
i: [1,78,0,12,5,0,0,1] dt_k1_funct_4
s: 78
d: [2,172,1,6,2,1,40,72,95]
s: 172
i: [1,79,0,12,5,0,0,1] dt_k1_funct_4
s: 79
i: [1,67,0,13,5,0,0,1] t25_funct_4
s: 67
d: [2,187,2,5,1,1,40,172,165]
a: [3,191,2,5,1,187]
s: 191
i: [1,77,0,17,5,0,0,255] rd3_funct_4
s: 77
d: [2,183,3,8,2,1,40,179,72]
s: 183
i: [1,68,0,23,9,0,0,1] t87_funct_4
s: 68
i: [1,69,0,25,9,0,0,1] t117_funct_4
s: 69
d: [2,216,1,26,8,1,40,68,85]
d: [2,225,1,24,7,2,21,216,125]
d: [2,226,1,22,6,3,21,225,130]
d: [2,227,1,20,5,4,21,226,135]
d: [2,228,1,18,4,5,21,227,140]
a: [3,232,1,4,1,228]
s: 232
d: [2,272,2,0,0,5,52,120,115,110,105,232,78]
e: 272
d: [2,273,2,8,4,1,40,232,78]
d: [2,275,2,6,3,2,21,273,105]
d: [2,276,2,4,2,3,21,275,110]
d: [2,277,2,2,1,4,21,276,115]
d: [2,278,2,0,0,5,21,277,120]
e: 278
r: 232
a: [3,231,1,4,1,228]
s: 231
a: [3,236,1,4,1,228]
s: 236
d: [2,280,2,0,0,5,52,120,115,110,105,236,79]
e: 280
d: [2,281,2,8,4,1,40,236,79]
d: [2,283,2,6,3,2,21,281,105]
d: [2,284,2,4,2,3,21,283,110]
d: [2,285,2,2,1,4,21,284,115]
d: [2,286,2,0,0,5,21,285,120]
e: 286
r: 236
a: [3,235,1,4,1,228]
s: 235
a: [3,240,1,5,1,228]
s: 240
d: [2,288,2,0,0,9,52,140,135,110,105,120,115,100,90,240,69]
e: 288
r: 240
d: [2,220,4,31,9,1,40,68,183]
a: [3,257,4,4,1,220]
s: 257
d: [2,290,5,0,0,3,52,140,135,73,257,66]
e: 290
r: 257
a: [3,239,1,5,1,228]
s: 239
a: [3,253,4,4,1,220]
s: 253
d: [2,293,5,0,0,2,52,140,73,253,65]
e: 293
r: 253
a: [3,244,1,5,1,228]
s: 244
d: [2,262,1,20,8,1,40,69,90]
d: [2,264,1,18,7,2,21,262,105]
d: [2,265,1,16,6,3,21,264,110]
d: [2,266,1,14,5,4,21,265,135]
d: [2,267,1,12,4,5,21,266,140]
s: 267
a: [3,252,4,4,1,220]
s: 252
d: [2,168,1,6,2,0,40,65,73]
s: 168
a: [3,256,4,4,1,220]
s: 256
d: [2,169,1,7,3,0,40,65,71]
s: 169
d: [2,181,1,8,3,0,40,66,73]
s: 181
d: [2,308,3,6,2,1,40,169,191]
d: [2,310,3,4,1,2,21,308,110]
a: [3,314,3,4,1,310]
s: 314
d: [2,182,1,9,4,0,40,66,71]
s: 182
d: [2,325,3,8,3,1,40,182,191]
d: [2,327,3,6,2,2,21,325,105]
d: [2,328,3,4,1,3,21,327,110]
a: [3,332,3,4,1,328]
s: 332
d: [2,195,1,16,6,0,40,67,72]
s: 195
d: [2,334,2,0,0,6,52,120,115,110,105,95,244,195]
e: 334
d: [2,339,2,11,5,1,40,195,244]
d: [2,349,2,9,4,2,21,339,120]
d: [2,350,2,7,3,3,21,349,105]
d: [2,351,2,5,2,4,21,350,110]
d: [2,352,2,3,1,5,21,351,115]
d: [2,353,2,0,0,6,21,352,95]
e: 353
f: 353,288,286,278,228,53,54,55,56,57,58,59,60,61,62,63,64
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t122_funct_4
% SZS output start Proof for funct_4__t122_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_funct_1(X3) & v1_relat_1(X3)) => ((r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X3)) & r1_tarski(X1,X3) & r1_tarski(X0,X2)) => r1_tarski(k1_funct_4(X0,X1),k1_funct_4(X2,X3))))))) [input t122_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_funct_1(X3) & v1_relat_1(X3)) => ((r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X3)) & r1_tarski(X1,X3) & r1_tarski(X0,X2)) => r1_tarski(k1_funct_4(X0,X1),k1_funct_4(X2,X3))))))) [negated conjecture 1]
5. ! [X0,X1] : ((v1_funct_1(X1) & v1_relat_1(X1) & v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k1_funct_4(X0,X1)) & v1_relat_1(k1_funct_4(X0,X1)))) [input dt_k1_funct_4]
15. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_tarski(X0,X2) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) => r1_tarski(X0,k1_funct_4(X2,X1)))))) [input t117_funct_4]
16. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
17. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => r1_tarski(X0,k1_funct_4(X1,X0)))) [input t25_funct_4]
19. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r1_tarski(X1,X2) & r1_tarski(X0,X2)) => r1_tarski(k1_funct_4(X0,X1),X2))))) [input t87_funct_4]
21. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 16]
23. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_tarski(k1_funct_4(X0,X1),k1_funct_4(X2,X3)) & (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X3)) & r1_tarski(X1,X3) & r1_tarski(X0,X2))) & (v1_funct_1(X3) & v1_relat_1(X3))) & (v1_funct_1(X2) & v1_relat_1(X2))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
24. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_tarski(k1_funct_4(X0,X1),k1_funct_4(X2,X3)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X3)) & r1_tarski(X1,X3) & r1_tarski(X0,X2) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 23]
28. ! [X0] : (! [X1] : (r1_tarski(X0,k1_funct_4(X1,X0)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 17]
29. ! [X0] : (! [X1] : (r1_tarski(X0,k1_funct_4(X1,X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 28]
30. ! [X0] : (! [X1] : (! [X2] : ((r1_tarski(k1_funct_4(X0,X1),X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X2))) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 19]
31. ! [X0] : (! [X1] : (! [X2] : (r1_tarski(k1_funct_4(X0,X1),X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 30]
32. ! [X0] : (! [X1] : (! [X2] : ((r1_tarski(X0,k1_funct_4(X2,X1)) | (~r1_tarski(X0,X2) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)))) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 15]
33. ! [X0] : (! [X1] : (! [X2] : (r1_tarski(X0,k1_funct_4(X2,X1)) | ~r1_tarski(X0,X2) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 32]
34. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 21]
35. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 34]
41. ! [X0,X1] : ((v1_funct_1(k1_funct_4(X0,X1)) & v1_relat_1(k1_funct_4(X0,X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 5]
42. ! [X0,X1] : ((v1_funct_1(k1_funct_4(X0,X1)) & v1_relat_1(k1_funct_4(X0,X1))) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 41]
43. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_tarski(k1_funct_4(X0,X1),k1_funct_4(X2,X3)) & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X3)) & r1_tarski(X1,X3) & r1_tarski(X0,X2) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r1_tarski(k1_funct_4(sK0,X1),k1_funct_4(X2,X3)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X3)) & r1_tarski(X1,X3) & r1_tarski(sK0,X2) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
44. ? [X1] : (? [X2] : (? [X3] : (~r1_tarski(k1_funct_4(sK0,X1),k1_funct_4(X2,X3)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X3)) & r1_tarski(X1,X3) & r1_tarski(sK0,X2) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(X1) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (~r1_tarski(k1_funct_4(sK0,sK1),k1_funct_4(X2,X3)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X3)) & r1_tarski(sK1,X3) & r1_tarski(sK0,X2) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
45. ? [X2] : (? [X3] : (~r1_tarski(k1_funct_4(sK0,sK1),k1_funct_4(X2,X3)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X3)) & r1_tarski(sK1,X3) & r1_tarski(sK0,X2) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(X2) & v1_relat_1(X2)) => (? [X3] : (~r1_tarski(k1_funct_4(sK0,sK1),k1_funct_4(sK2,X3)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X3)) & r1_tarski(sK1,X3) & r1_tarski(sK0,sK2) & v1_funct_1(X3) & v1_relat_1(X3)) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
46. ? [X3] : (~r1_tarski(k1_funct_4(sK0,sK1),k1_funct_4(sK2,X3)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X3)) & r1_tarski(sK1,X3) & r1_tarski(sK0,sK2) & v1_funct_1(X3) & v1_relat_1(X3)) => (~r1_tarski(k1_funct_4(sK0,sK1),k1_funct_4(sK2,sK3)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK3)) & r1_tarski(sK1,sK3) & r1_tarski(sK0,sK2) & v1_funct_1(sK3) & v1_relat_1(sK3)) [choice axiom]
47. (((~r1_tarski(k1_funct_4(sK0,sK1),k1_funct_4(sK2,sK3)) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK3)) & r1_tarski(sK1,sK3) & r1_tarski(sK0,sK2) & v1_funct_1(sK3) & v1_relat_1(sK3)) & v1_funct_1(sK2) & v1_relat_1(sK2)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 24,46,45,44,43]
53. v1_relat_1(sK0) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
54. v1_funct_1(sK0) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
55. v1_relat_1(sK1) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
56. v1_funct_1(sK1) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
57. v1_relat_1(sK2) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
58. v1_funct_1(sK2) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
59. v1_relat_1(sK3) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
60. v1_funct_1(sK3) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
61. r1_tarski(sK0,sK2) [cnf transformation 47] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
62. r1_tarski(sK1,sK3) [cnf transformation 47] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
63. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK3)) [cnf transformation 47] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
64. ~r1_tarski(k1_funct_4(sK0,sK1),k1_funct_4(sK2,sK3)) [cnf transformation 47] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
67. r1_tarski(X0,k1_funct_4(X1,X0)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. r1_tarski(k1_funct_4(X0,X1),X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 31] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X2) | r1_tarski(X0,k1_funct_4(X2,X1)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 35] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
78. v1_relat_1(k1_funct_4(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 42] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. v1_funct_1(k1_funct_4(X0,X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 42] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
83. 1 <=> r1_tarski(k1_funct_4(sK0,sK1),k1_funct_4(sK2,sK3)) [avatar definition]
85. ~r1_tarski(k1_funct_4(sK0,sK1),k1_funct_4(sK2,sK3)) <- (~1) [avatar component clause 83] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
86. ~1 [avatar split clause 64,83]
88. 2 <=> r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK3)) [avatar definition]
90. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK3)) <- (2) [avatar component clause 88] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
91. 2 [avatar split clause 63,88]
93. 3 <=> r1_tarski(sK1,sK3) [avatar definition]
95. r1_tarski(sK1,sK3) <- (3) [avatar component clause 93] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
96. 3 [avatar split clause 62,93]
98. 4 <=> r1_tarski(sK0,sK2) [avatar definition]
100. r1_tarski(sK0,sK2) <- (4) [avatar component clause 98] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
101. 4 [avatar split clause 61,98]
103. 5 <=> v1_funct_1(sK3) [avatar definition]
105. v1_funct_1(sK3) <- (5) [avatar component clause 103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
106. 5 [avatar split clause 60,103]
108. 6 <=> v1_relat_1(sK3) [avatar definition]
110. v1_relat_1(sK3) <- (6) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
111. 6 [avatar split clause 59,108]
113. 7 <=> v1_funct_1(sK2) [avatar definition]
115. v1_funct_1(sK2) <- (7) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
116. 7 [avatar split clause 58,113]
118. 8 <=> v1_relat_1(sK2) [avatar definition]
120. v1_relat_1(sK2) <- (8) [avatar component clause 118] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
121. 8 [avatar split clause 57,118]
123. 9 <=> v1_funct_1(sK1) [avatar definition]
125. v1_funct_1(sK1) <- (9) [avatar component clause 123] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
126. 9 [avatar split clause 56,123]
128. 10 <=> v1_relat_1(sK1) [avatar definition]
130. v1_relat_1(sK1) <- (10) [avatar component clause 128] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
131. 10 [avatar split clause 55,128]
133. 11 <=> v1_funct_1(sK0) [avatar definition]
135. v1_funct_1(sK0) <- (11) [avatar component clause 133] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
136. 11 [avatar split clause 54,133]
138. 12 <=> v1_relat_1(sK0) [avatar definition]
140. v1_relat_1(sK0) <- (12) [avatar component clause 138] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
141. 12 [avatar split clause 53,138]
195. r1_tarski(X2,k1_funct_4(X0,X1)) | ~v1_relat_1(X0) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~r1_tarski(X2,X1) [resolution 67,72] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
216. ~r1_tarski(sK1,k1_funct_4(sK2,sK3)) | ~r1_tarski(sK0,k1_funct_4(sK2,sK3)) | ~v1_funct_1(k1_funct_4(sK2,sK3)) | ~v1_relat_1(k1_funct_4(sK2,sK3)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [resolution 68,85] {a:1,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
225. ~r1_tarski(sK1,k1_funct_4(sK2,sK3)) | ~r1_tarski(sK0,k1_funct_4(sK2,sK3)) | ~v1_funct_1(k1_funct_4(sK2,sK3)) | ~v1_relat_1(k1_funct_4(sK2,sK3)) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 9) [subsumption resolution 216,125] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
226. ~r1_tarski(sK1,k1_funct_4(sK2,sK3)) | ~r1_tarski(sK0,k1_funct_4(sK2,sK3)) | ~v1_funct_1(k1_funct_4(sK2,sK3)) | ~v1_relat_1(k1_funct_4(sK2,sK3)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 9, 10) [subsumption resolution 225,130] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
227. ~r1_tarski(sK1,k1_funct_4(sK2,sK3)) | ~r1_tarski(sK0,k1_funct_4(sK2,sK3)) | ~v1_funct_1(k1_funct_4(sK2,sK3)) | ~v1_relat_1(k1_funct_4(sK2,sK3)) | ~v1_relat_1(sK0) <- (~1, 9, 10, 11) [subsumption resolution 226,135] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
228. ~r1_tarski(sK1,k1_funct_4(sK2,sK3)) | ~r1_tarski(sK0,k1_funct_4(sK2,sK3)) | ~v1_funct_1(k1_funct_4(sK2,sK3)) | ~v1_relat_1(k1_funct_4(sK2,sK3)) <- (~1, 9, 10, 11, 12) [subsumption resolution 227,140] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
230. 21 <=> v1_relat_1(k1_funct_4(sK2,sK3)) [avatar definition]
232. ~v1_relat_1(k1_funct_4(sK2,sK3)) <- (~21) [avatar component clause 230] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
234. 22 <=> v1_funct_1(k1_funct_4(sK2,sK3)) [avatar definition]
236. ~v1_funct_1(k1_funct_4(sK2,sK3)) <- (~22) [avatar component clause 234] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
238. 23 <=> r1_tarski(sK0,k1_funct_4(sK2,sK3)) [avatar definition]
240. ~r1_tarski(sK0,k1_funct_4(sK2,sK3)) <- (~23) [avatar component clause 238] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
242. 24 <=> r1_tarski(sK1,k1_funct_4(sK2,sK3)) [avatar definition]
244. ~r1_tarski(sK1,k1_funct_4(sK2,sK3)) <- (~24) [avatar component clause 242] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
245. ~21 | ~22 | ~23 | ~24 | 1 | ~9 | ~10 | ~11 | ~12 [avatar split clause 228,138,133,128,123,83,242,238,234,230]
273. ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (~21) [resolution 232,78] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
275. ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (5, ~21) [subsumption resolution 273,105] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
276. ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (5, 6, ~21) [subsumption resolution 275,110] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
277. ~v1_relat_1(sK2) <- (5, 6, 7, ~21) [subsumption resolution 276,115] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
278. $false <- (5, 6, 7, 8, ~21) [subsumption resolution 277,120] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
279. ~5 | ~6 | ~7 | ~8 | 21 [avatar contradiction clause 278]
281. ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (~22) [resolution 236,79] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
283. ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (5, ~22) [subsumption resolution 281,105] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
284. ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (5, 6, ~22) [subsumption resolution 283,110] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
285. ~v1_relat_1(sK2) <- (5, 6, 7, ~22) [subsumption resolution 284,115] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
286. $false <- (5, 6, 7, 8, ~22) [subsumption resolution 285,120] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
287. ~5 | ~6 | ~7 | ~8 | 22 [avatar contradiction clause 286]
288. $false <- (2, 4, 5, 6, 7, 8, 11, 12, ~23) [unit resulting resolution 140,135,110,105,120,115,100,90,240,69] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:47,thDist:-47}
289. ~2 | ~4 | ~5 | ~6 | ~7 | ~8 | ~11 | ~12 | 23 [avatar contradiction clause 288]
339. ~v1_relat_1(sK2) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~r1_tarski(sK1,sK3) <- (~24) [resolution 195,244] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
349. ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~r1_tarski(sK1,sK3) <- (8, ~24) [subsumption resolution 339,120] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
350. ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~r1_tarski(sK1,sK3) <- (5, 8, ~24) [subsumption resolution 349,105] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
351. ~v1_funct_1(sK2) | ~r1_tarski(sK1,sK3) <- (5, 6, 8, ~24) [subsumption resolution 350,110] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
352. ~r1_tarski(sK1,sK3) <- (5, 6, 7, 8, ~24) [subsumption resolution 351,115] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
353. $false <- (3, 5, 6, 7, 8, ~24) [subsumption resolution 352,95] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
354. ~3 | ~5 | ~6 | ~7 | ~8 | 24 [avatar contradiction clause 353]
355. $false [avatar sat refutation 86,91,96,101,106,111,116,121,126,131,136,141,245,279,287,289,354] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:384,thDist:-384}
% SZS output end Proof for funct_4__t122_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 190
% Active clauses: 55
% Passive clauses: 65
% Final active clauses: 50
% Final passive clauses: 10
% Main loop iterations started: 55
% 
% >>> Simplifying Inferences
% Duplicate literals: 42
% Fw subsumption resolutions: 36
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 25
% 
% >>> Generating Inferences
% Binary resolution: 57
% Unit resulting resolution: 6
% Forward superposition: 10
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 9
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 34
% SAT solver unit clauses: 14
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.003 s

