i: [1,66,0,2,1,1,0,0] t12_conmetr1
a: [3,92,0,2,1,66]
s: 92
i: [1,65,0,2,1,1,0,0] t12_conmetr1
a: [3,97,0,2,1,65]
s: 97
i: [1,64,0,2,1,1,0,0] t12_conmetr1
a: [3,102,0,2,1,64]
s: 102
i: [1,63,0,2,1,1,0,0] t12_conmetr1
a: [3,107,0,2,1,63]
s: 107
i: [1,62,0,2,1,1,0,0] t12_conmetr1
a: [3,112,0,2,1,62]
s: 112
i: [1,61,0,2,1,1,0,0] t12_conmetr1
a: [3,117,0,2,1,61]
s: 117
i: [1,87,0,2,1,0,0,1] existence_l1_analoaf
a: [3,122,0,2,1,87]
s: 122
i: [1,88,0,2,1,0,0,2] existence_l1_struct_0
a: [3,127,0,2,1,88]
s: 127
i: [1,67,0,4,2,0,0,1] dt_l1_analoaf
s: 67
i: [1,68,0,6,3,0,0,255] cc7_struct_0
s: 68
i: [1,69,0,6,3,0,0,255] cc5_struct_0
s: 69
d: [2,129,1,2,1,1,40,67,102]
a: [3,134,1,2,1,129]
s: 134
i: [1,71,0,6,3,0,0,255] cc17_struct_0
s: 71
d: [2,130,1,2,1,1,40,67,122]
a: [3,139,1,2,1,130]
s: 139
i: [1,72,0,6,3,0,0,255] cc2_struct_0
s: 72
i: [1,73,0,6,3,0,0,255] cc1_struct_0
s: 73
d: [2,141,1,4,2,1,40,68,127]
a: [3,145,1,2,1,141]
s: 145
i: [1,76,0,6,3,0,0,255] cc4_struct_0
s: 76
d: [2,152,2,4,2,1,40,139,68]
a: [3,156,2,2,1,152]
s: 156
i: [1,77,0,6,3,0,0,255] cc6_struct_0
s: 77
i: [1,74,0,7,3,0,0,255] cc8_struct_0
s: 74
i: [1,78,0,7,3,0,0,255] cc9_struct_0
s: 78
i: [1,79,0,7,3,0,0,255] cc11_struct_0
s: 79
i: [1,80,0,7,3,0,0,255] cc11_struct_0
s: 80
i: [1,81,0,9,4,0,0,255] cc10_struct_0
s: 81
i: [1,82,0,12,6,0,0,2] t9_aff_2
s: 82
i: [1,83,0,12,6,0,0,1] t10_conmetr1
s: 83
i: [1,84,0,12,6,0,0,1] t10_conmetr1
s: 84
d: [2,168,1,10,5,1,40,84,97]
d: [2,169,1,8,4,2,21,168,102]
d: [2,170,1,6,3,3,21,169,107]
d: [2,171,1,4,2,4,21,170,112]
d: [2,172,1,2,1,5,21,171,117]
a: [3,176,1,2,1,172]
s: 176
i: [1,85,0,12,6,0,0,1] t9_conmetr1
s: 85
i: [1,86,0,12,6,0,0,1] t9_conmetr1
s: 86
d: [2,179,2,10,5,1,40,176,82]
d: [2,180,2,8,4,2,21,179,102]
d: [2,181,2,6,3,3,21,180,107]
d: [2,182,2,4,2,4,21,181,112]
d: [2,183,2,2,1,5,21,182,117]
a: [3,187,2,2,1,183]
s: 187
d: [2,189,3,0,0,6,52,117,112,107,102,92,187,85]
e: 189
d: [2,190,3,10,5,1,40,187,85]
d: [2,192,3,8,4,2,21,190,92]
d: [2,193,3,6,3,3,21,192,102]
d: [2,194,3,4,2,4,21,193,107]
d: [2,195,3,2,1,5,21,194,112]
d: [2,196,3,0,0,6,21,195,117]
e: 196
f: 196,183,172,61,62,63,64,65,66
% Refutation found. Thanks to Tanya!
% SZS status Theorem for conmetr1__t12_conmetr1
% SZS output start Proof for conmetr1__t12_conmetr1
1. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v6_conmetr1(X0) => v5_conmetr1(X0))) [input t12_conmetr1]
2. ~! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v6_conmetr1(X0) => v5_conmetr1(X0))) [negated conjecture 1]
19. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v2_aff_2(X0) <=> v6_conmetr1(X0))) [input t10_conmetr1]
20. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v2_aff_2(X0) => v13_aff_2(X0))) [input t9_aff_2]
21. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v13_aff_2(X0) <=> v5_conmetr1(X0))) [input t9_conmetr1]
22. ? [X0] : ((~v5_conmetr1(X0) & v6_conmetr1(X0)) & (l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
23. ? [X0] : (~v5_conmetr1(X0) & v6_conmetr1(X0) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) [flattening 22]
47. ! [X0] : ((v13_aff_2(X0) | ~v2_aff_2(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 20]
48. ! [X0] : (v13_aff_2(X0) | ~v2_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 47]
49. ! [X0] : ((v2_aff_2(X0) <=> v6_conmetr1(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 19]
50. ! [X0] : ((v2_aff_2(X0) <=> v6_conmetr1(X0)) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 49]
51. ! [X0] : ((v13_aff_2(X0) <=> v5_conmetr1(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 21]
52. ! [X0] : ((v13_aff_2(X0) <=> v5_conmetr1(X0)) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 51]
53. ? [X0] : (~v5_conmetr1(X0) & v6_conmetr1(X0) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (~v5_conmetr1(sK0) & v6_conmetr1(sK0) & l1_analoaf(sK0) & v2_diraf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
54. ~v5_conmetr1(sK0) & v6_conmetr1(sK0) & l1_analoaf(sK0) & v2_diraf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0) [skolemisation 23,53]
55. ! [X0] : (((v2_aff_2(X0) | ~v6_conmetr1(X0)) & (v6_conmetr1(X0) | ~v2_aff_2(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [nnf transformation 50]
56. ! [X0] : (((v13_aff_2(X0) | ~v5_conmetr1(X0)) & (v5_conmetr1(X0) | ~v13_aff_2(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [nnf transformation 52]
61. ~v7_struct_0(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
62. v1_diraf(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
63. v2_diraf(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
64. l1_analoaf(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. v6_conmetr1(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
66. ~v5_conmetr1(sK0) [cnf transformation 54] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
82. ~v2_aff_2(X0) | v13_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 48] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
84. ~v6_conmetr1(X0) | v2_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 55] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
85. ~v13_aff_2(X0) | v5_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 56] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
90. 1 <=> v5_conmetr1(sK0) [avatar definition]
92. ~v5_conmetr1(sK0) <- (~1) [avatar component clause 90] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
93. ~1 [avatar split clause 66,90]
95. 2 <=> v6_conmetr1(sK0) [avatar definition]
97. v6_conmetr1(sK0) <- (2) [avatar component clause 95] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
98. 2 [avatar split clause 65,95]
100. 3 <=> l1_analoaf(sK0) [avatar definition]
102. l1_analoaf(sK0) <- (3) [avatar component clause 100] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
103. 3 [avatar split clause 64,100]
105. 4 <=> v2_diraf(sK0) [avatar definition]
107. v2_diraf(sK0) <- (4) [avatar component clause 105] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
108. 4 [avatar split clause 63,105]
110. 5 <=> v1_diraf(sK0) [avatar definition]
112. v1_diraf(sK0) <- (5) [avatar component clause 110] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
113. 5 [avatar split clause 62,110]
115. 6 <=> v7_struct_0(sK0) [avatar definition]
117. ~v7_struct_0(sK0) <- (~6) [avatar component clause 115] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
118. ~6 [avatar split clause 61,115]
168. v2_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2) [resolution 84,97] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
169. v2_aff_2(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3) [subsumption resolution 168,102] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
170. v2_aff_2(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3, 4) [subsumption resolution 169,107] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
171. v2_aff_2(sK0) | v7_struct_0(sK0) <- (2, 3, 4, 5) [subsumption resolution 170,112] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
172. v2_aff_2(sK0) <- (2, 3, 4, 5, ~6) [subsumption resolution 171,117] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
174. 15 <=> v2_aff_2(sK0) [avatar definition]
176. v2_aff_2(sK0) <- (15) [avatar component clause 174] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
177. 15 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 172,115,110,105,100,95,174]
179. v13_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (15) [resolution 176,82] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
180. v13_aff_2(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 15) [subsumption resolution 179,102] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
181. v13_aff_2(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 4, 15) [subsumption resolution 180,107] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
182. v13_aff_2(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 15) [subsumption resolution 181,112] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
183. v13_aff_2(sK0) <- (3, 4, 5, ~6, 15) [subsumption resolution 182,117] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
185. 16 <=> v13_aff_2(sK0) [avatar definition]
187. v13_aff_2(sK0) <- (16) [avatar component clause 185] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
188. 16 | ~3 | ~4 | ~5 | 6 | ~15 [avatar split clause 183,174,115,110,105,100,185]
190. v5_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (16) [resolution 187,85] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
192. ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 16) [subsumption resolution 190,92] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
193. ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 3, 16) [subsumption resolution 192,102] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
194. ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~1, 3, 4, 16) [subsumption resolution 193,107] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
195. v7_struct_0(sK0) <- (~1, 3, 4, 5, 16) [subsumption resolution 194,112] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
196. $false <- (~1, 3, 4, 5, ~6, 16) [subsumption resolution 195,117] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
197. 1 | ~3 | ~4 | ~5 | 6 | ~16 [avatar contradiction clause 196]
198. $false [avatar sat refutation 93,98,103,108,113,118,177,188,197] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:86,thDist:-86}
% SZS output end Proof for conmetr1__t12_conmetr1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 28
% Generated clauses: 71
% Active clauses: 28
% Passive clauses: 32
% Final active clauses: 28
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 12
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 8
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

