i: [1,55,0,5,2,1,0,0] t3_projpl_1
a: [3,79,0,2,1,55]
s: 79
i: [1,46,0,2,1,1,0,0] t3_projpl_1
a: [3,127,0,2,1,46]
s: 127
i: [1,73,0,2,1,0,0,2] rc2_xboole_0
a: [3,132,0,2,1,73]
s: 132
i: [1,74,0,2,1,0,0,2] rc1_xboole_0
a: [3,137,0,2,1,74]
s: 137
i: [1,75,0,2,1,0,0,1] existence_l1_incsp_1
a: [3,142,0,2,1,75]
s: 142
i: [1,72,0,4,1,0,0,1] existence_m1_subset_1
s: 72
i: [1,56,0,5,2,0,0,1] dt_u1_incsp_1
s: 56
i: [1,57,0,5,2,0,0,1] dt_u2_incsp_1
s: 57
i: [1,47,0,34,9,1,0,0] t3_projpl_1
a: [3,122,0,32,8,47]
s: 122
i: [1,59,0,9,3,0,0,1] d4_projpl_1
s: 59
i: [1,60,0,9,3,0,0,1] d4_projpl_1
s: 60
i: [1,61,0,9,3,0,0,1] d4_projpl_1
s: 61
i: [1,62,0,9,3,0,0,1] d4_projpl_1
s: 62
i: [1,67,0,9,3,0,0,1] d4_projpl_1
s: 67
i: [1,68,0,9,3,0,0,1] d4_projpl_1
s: 68
i: [1,63,0,10,3,0,0,1] d4_projpl_1
s: 63
i: [1,64,0,10,3,0,0,1] d4_projpl_1
s: 64
i: [1,65,0,10,3,0,0,1] d4_projpl_1
s: 65
i: [1,66,0,10,3,0,0,1] d4_projpl_1
s: 66
i: [1,69,0,23,6,0,0,1] d2_projpl_1
s: 69
i: [1,70,0,23,6,0,0,1] d2_projpl_1
s: 70
i: [1,71,0,27,7,0,0,1] d2_projpl_1
s: 71
i: [1,58,0,42,12,0,0,1] d4_projpl_1
s: 58
d: [2,144,1,49,13,1,40,71,122]
d: [2,149,1,37,10,1,18,144]
d: [2,150,1,35,9,2,21,149,127]
s: 150
d: [2,159,2,52,14,2,40,150,71]
d: [2,160,2,40,11,2,18,159]
d: [2,161,2,38,10,2,21,160,127]
s: 161
d: [2,162,3,44,11,2,40,161,63]
d: [2,166,3,42,10,3,21,162,79]
d: [2,167,3,40,9,3,21,166,127]
a: [3,171,3,5,1,167]
s: 171
d: [2,206,4,0,0,3,52,127,79,171,59]
e: 206
d: [2,207,4,4,2,1,40,171,59]
d: [2,209,4,2,1,2,21,207,79]
d: [2,210,4,0,0,3,21,209,127]
e: 210
r: 79
a: [3,78,0,2,1,55]
s: 78
a: [3,83,0,3,1,55]
s: 83
i: [1,54,0,5,2,1,0,0] t3_projpl_1
a: [3,88,0,3,1,54]
s: 88
i: [1,51,0,6,2,1,0,0] t3_projpl_1
a: [3,103,0,4,1,51]
s: 103
i: [1,50,0,6,2,1,0,0] t3_projpl_1
a: [3,108,0,4,1,50]
s: 108
i: [1,49,0,6,2,1,0,0] t3_projpl_1
a: [3,113,0,4,1,49]
s: 113
i: [1,48,0,6,2,1,0,0] t3_projpl_1
a: [3,118,0,4,1,48]
s: 118
i: [1,53,0,7,2,1,0,0] t3_projpl_1
a: [3,93,0,5,1,53]
s: 93
r: 78
r: 108
r: 171
s: 79
a: [3,107,0,4,1,50]
s: 107
i: [1,52,0,7,2,1,0,0] t3_projpl_1
a: [3,98,0,5,1,52]
s: 98
a: [3,170,3,5,1,167]
s: 170
a: [3,178,3,5,1,167]
s: 178
d: [2,274,4,0,0,3,52,127,79,178,61]
e: 274
d: [2,275,4,4,2,1,40,178,61]
d: [2,277,4,2,1,2,21,275,79]
d: [2,278,4,0,0,3,21,277,127]
e: 278
r: 178
d: [2,163,3,44,11,2,40,161,64]
d: [2,180,3,42,10,3,21,163,79]
d: [2,181,3,40,9,3,21,180,127]
a: [3,185,3,5,1,181]
s: 185
d: [2,280,4,0,0,3,52,127,79,185,60]
e: 280
d: [2,281,4,4,2,1,40,185,60]
d: [2,283,4,2,1,2,21,281,79]
d: [2,284,4,0,0,3,21,283,127]
e: 284
r: 185
d: [2,164,3,44,11,2,40,161,65]
d: [2,190,3,42,10,3,21,164,79]
d: [2,191,3,40,9,3,21,190,127]
a: [3,198,3,5,1,191]
s: 198
d: [2,286,4,0,0,3,52,127,79,198,62]
e: 286
d: [2,287,4,4,2,1,40,198,62]
d: [2,289,4,2,1,2,21,287,79]
d: [2,290,4,0,0,3,21,289,127]
e: 290
r: 198
a: [3,174,3,30,7,167]
s: 174
a: [3,188,3,30,7,181]
s: 188
d: [2,293,4,32,8,1,40,174,65]
d: [2,295,4,30,7,2,21,293,79]
d: [2,296,4,28,6,3,21,295,127]
a: [3,303,4,5,1,296]
s: 303
d: [2,314,5,0,0,3,52,127,79,303,68]
e: 314
d: [2,317,5,9,3,1,44,68,303]
d: [2,319,5,4,2,1,19,317]
d: [2,321,5,2,1,2,21,319,79]
d: [2,322,5,0,0,3,21,321,127]
e: 322
r: 303
a: [3,194,3,30,7,191]
s: 194
d: [2,165,3,44,11,2,40,161,66]
d: [2,200,3,42,10,3,21,165,79]
d: [2,201,3,40,9,3,21,200,127]
a: [3,204,3,30,7,201]
s: 204
a: [3,299,4,18,4,296]
s: 299
d: [2,306,4,32,8,1,40,188,66]
d: [2,308,4,30,7,2,21,306,79]
d: [2,309,4,28,6,3,21,308,127]
a: [3,312,4,18,4,309]
s: 312
d: [2,337,5,20,5,1,40,299,66]
d: [2,338,5,18,4,2,21,337,79]
d: [2,339,5,16,3,3,21,338,127]
a: [3,347,5,6,1,339]
s: 347
d: [2,359,6,0,0,3,52,127,79,347,64]
e: 359
d: [2,360,6,4,2,1,40,347,64]
d: [2,362,6,2,1,2,21,360,79]
d: [2,363,6,0,0,3,21,362,127]
e: 363
r: 347
a: [3,343,5,5,1,339]
s: 343
r: 312
r: 204
r: 188
d: [2,368,6,0,0,3,52,127,79,343,67]
e: 368
d: [2,371,6,9,3,1,44,67,343]
d: [2,373,6,4,2,1,19,371]
d: [2,377,6,2,1,2,21,373,79]
d: [2,378,6,0,0,3,21,377,127]
e: 378
r: 79
r: 107
r: 343
s: 78
s: 108
d: [2,418,1,0,0,9,52,88,83,113,103,93,98,118,108,122]
e: 418
d: [2,419,1,0,0,9,52,88,83,118,103,98,93,113,108,122]
e: 419
r: 161
r: 150
r: 122
d: [2,220,1,18,5,1,40,93,70]
d: [2,230,1,14,4,2,21,220,103]
d: [2,231,1,10,3,3,21,230,108]
d: [2,232,1,6,2,4,21,231,113]
d: [2,233,1,4,1,5,21,232,127]
a: [3,237,1,4,1,233]
s: 237
d: [2,221,1,18,5,1,40,93,69]
d: [2,239,1,14,4,2,21,221,103]
d: [2,240,1,10,3,3,21,239,108]
d: [2,241,1,6,2,4,21,240,113]
d: [2,242,1,4,1,5,21,241,127]
a: [3,246,1,4,1,242]
s: 246
d: [2,273,1,18,5,1,40,98,69]
d: [2,388,1,14,4,2,21,273,103]
d: [2,389,1,10,3,3,21,388,118]
d: [2,390,1,8,2,4,21,389,127]
a: [3,394,1,4,1,390]
s: 394
d: [2,272,1,18,5,1,40,98,70]
d: [2,396,1,14,4,2,21,272,103]
d: [2,397,1,10,3,3,21,396,118]
d: [2,398,1,8,2,4,21,397,127]
a: [3,402,1,4,1,398]
s: 402
d: [2,437,2,0,0,12,52,127,78,83,88,113,237,246,108,394,118,103,402,58]
e: 437
d: [2,438,2,0,0,12,52,127,78,88,83,113,246,237,108,394,118,103,402,58]
e: 438
d: [2,439,2,0,0,12,52,127,78,83,88,246,108,394,113,237,118,103,402,58]
e: 439
d: [2,440,2,0,0,12,52,127,78,88,83,394,108,246,113,237,118,103,402,58]
e: 440
d: [2,212,4,2,1,2,21,207,127]
d: [2,248,1,10,3,3,21,239,113]
d: [2,249,1,8,2,4,21,248,127]
d: [2,251,1,10,3,3,21,230,113]
d: [2,252,1,8,2,4,21,251,127]
d: [2,268,3,42,10,2,21,164,127]
d: [2,324,4,32,8,1,40,194,63]
d: [2,327,4,30,7,2,21,324,79]
d: [2,328,4,28,6,3,21,327,127]
d: [2,382,5,18,4,2,21,337,127]
f: 440,398,390,382,378,363,328,322,290,284,278,268,252,249,212,46,47,48,49,50,51,52,53,54,55
% Refutation found. Thanks to Tanya!
% SZS status Theorem for projpl_1__t3_projpl_1
% SZS output start Proof for projpl_1__t3_projpl_1
1. ! [X0] : (l1_incsp_1(X0) => (v1_projpl_1(X0) <=> ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => ! [X4] : (m1_subset_1(X4,u2_incsp_1(X0)) => ~(X3 != X4 & X1 != X2 & r2_projpl_1(X0,X2,X3,X4) & r2_projpl_1(X0,X1,X3,X4)))))))) [input t3_projpl_1]
2. ~! [X0] : (l1_incsp_1(X0) => (v1_projpl_1(X0) <=> ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => ! [X4] : (m1_subset_1(X4,u2_incsp_1(X0)) => ~(X3 != X4 & X1 != X2 & r2_projpl_1(X0,X2,X3,X4) & r2_projpl_1(X0,X1,X3,X4)))))))) [negated conjecture 1]
3. ! [X0] : (l1_incsp_1(X0) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u2_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => (r2_projpl_1(X0,X1,X2,X3) <=> (r1_incsp_1(X0,X1,X3) & r1_incsp_1(X0,X1,X2))))))) [input d2_projpl_1]
4. ! [X0] : (l1_incsp_1(X0) => (v1_projpl_1(X0) <=> ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => ! [X4] : (m1_subset_1(X4,u2_incsp_1(X0)) => ~(X3 != X4 & X1 != X2 & r1_incsp_1(X0,X2,X4) & r1_incsp_1(X0,X1,X4) & r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,X1,X3)))))))) [input d4_projpl_1]
13. ? [X0] : ((v1_projpl_1(X0) <~> ! [X1] : (! [X2] : (! [X3] : (! [X4] : ((X3 = X4 | X1 = X2 | ~r2_projpl_1(X0,X2,X3,X4) | ~r2_projpl_1(X0,X1,X3,X4)) | ~m1_subset_1(X4,u2_incsp_1(X0))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0)))) & l1_incsp_1(X0)) [ennf transformation 2]
14. ? [X0] : ((v1_projpl_1(X0) <~> ! [X1] : (! [X2] : (! [X3] : (! [X4] : (X3 = X4 | X1 = X2 | ~r2_projpl_1(X0,X2,X3,X4) | ~r2_projpl_1(X0,X1,X3,X4) | ~m1_subset_1(X4,u2_incsp_1(X0))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0)))) & l1_incsp_1(X0)) [flattening 13]
17. ! [X0] : ((v1_projpl_1(X0) <=> ! [X1] : (! [X2] : (! [X3] : (! [X4] : ((X3 = X4 | X1 = X2 | ~r1_incsp_1(X0,X2,X4) | ~r1_incsp_1(X0,X1,X4) | ~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3)) | ~m1_subset_1(X4,u2_incsp_1(X0))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0)))) | ~l1_incsp_1(X0)) [ennf transformation 4]
18. ! [X0] : ((v1_projpl_1(X0) <=> ! [X1] : (! [X2] : (! [X3] : (! [X4] : (X3 = X4 | X1 = X2 | ~r1_incsp_1(X0,X2,X4) | ~r1_incsp_1(X0,X1,X4) | ~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3) | ~m1_subset_1(X4,u2_incsp_1(X0))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0)))) | ~l1_incsp_1(X0)) [flattening 17]
19. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_projpl_1(X0,X1,X2,X3) <=> (r1_incsp_1(X0,X1,X3) & r1_incsp_1(X0,X1,X2))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u2_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [ennf transformation 3]
20. ? [X0] : (((? [X1] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & X1 != X2 & r2_projpl_1(X0,X2,X3,X4) & r2_projpl_1(X0,X1,X3,X4) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) | ~v1_projpl_1(X0)) & (! [X1] : (! [X2] : (! [X3] : (! [X4] : (X3 = X4 | X1 = X2 | ~r2_projpl_1(X0,X2,X3,X4) | ~r2_projpl_1(X0,X1,X3,X4) | ~m1_subset_1(X4,u2_incsp_1(X0))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | v1_projpl_1(X0))) & l1_incsp_1(X0)) [nnf transformation 14]
21. ? [X0] : ((? [X1] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & X1 != X2 & r2_projpl_1(X0,X2,X3,X4) & r2_projpl_1(X0,X1,X3,X4) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) | ~v1_projpl_1(X0)) & (! [X1] : (! [X2] : (! [X3] : (! [X4] : (X3 = X4 | X1 = X2 | ~r2_projpl_1(X0,X2,X3,X4) | ~r2_projpl_1(X0,X1,X3,X4) | ~m1_subset_1(X4,u2_incsp_1(X0))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | v1_projpl_1(X0)) & l1_incsp_1(X0)) [flattening 20]
22. ? [X0] : ((? [X1] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & X1 != X2 & r2_projpl_1(X0,X2,X3,X4) & r2_projpl_1(X0,X1,X3,X4) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) | ~v1_projpl_1(X0)) & (! [X5] : (! [X6] : (! [X7] : (! [X8] : (X7 = X8 | X5 = X6 | ~r2_projpl_1(X0,X6,X7,X8) | ~r2_projpl_1(X0,X5,X7,X8) | ~m1_subset_1(X8,u2_incsp_1(X0))) | ~m1_subset_1(X7,u2_incsp_1(X0))) | ~m1_subset_1(X6,u1_incsp_1(X0))) | ~m1_subset_1(X5,u1_incsp_1(X0))) | v1_projpl_1(X0)) & l1_incsp_1(X0)) [rectify 21]
23. ? [X0] : ((? [X1] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & X1 != X2 & r2_projpl_1(X0,X2,X3,X4) & r2_projpl_1(X0,X1,X3,X4) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) | ~v1_projpl_1(X0)) & (! [X5] : (! [X6] : (! [X7] : (! [X8] : (X7 = X8 | X5 = X6 | ~r2_projpl_1(X0,X6,X7,X8) | ~r2_projpl_1(X0,X5,X7,X8) | ~m1_subset_1(X8,u2_incsp_1(X0))) | ~m1_subset_1(X7,u2_incsp_1(X0))) | ~m1_subset_1(X6,u1_incsp_1(X0))) | ~m1_subset_1(X5,u1_incsp_1(X0))) | v1_projpl_1(X0)) & l1_incsp_1(X0)) => ((? [X1] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & X1 != X2 & r2_projpl_1(sK0,X2,X3,X4) & r2_projpl_1(sK0,X1,X3,X4) & m1_subset_1(X4,u2_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) | ~v1_projpl_1(sK0)) & (! [X5] : (! [X6] : (! [X7] : (! [X8] : (X7 = X8 | X5 = X6 | ~r2_projpl_1(sK0,X6,X7,X8) | ~r2_projpl_1(sK0,X5,X7,X8) | ~m1_subset_1(X8,u2_incsp_1(sK0))) | ~m1_subset_1(X7,u2_incsp_1(sK0))) | ~m1_subset_1(X6,u1_incsp_1(sK0))) | ~m1_subset_1(X5,u1_incsp_1(sK0))) | v1_projpl_1(sK0)) & l1_incsp_1(sK0)) [choice axiom]
24. ? [X1] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & X1 != X2 & r2_projpl_1(sK0,X2,X3,X4) & r2_projpl_1(sK0,X1,X3,X4) & m1_subset_1(X4,u2_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) => (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & sK1 != X2 & r2_projpl_1(sK0,X2,X3,X4) & r2_projpl_1(sK0,sK1,X3,X4) & m1_subset_1(X4,u2_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) [choice axiom]
25. ? [X2] : (? [X3] : (? [X4] : (X3 != X4 & sK1 != X2 & r2_projpl_1(sK0,X2,X3,X4) & r2_projpl_1(sK0,sK1,X3,X4) & m1_subset_1(X4,u2_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) => (? [X3] : (? [X4] : (X3 != X4 & sK1 != sK2 & r2_projpl_1(sK0,sK2,X3,X4) & r2_projpl_1(sK0,sK1,X3,X4) & m1_subset_1(X4,u2_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) [choice axiom]
26. ? [X3] : (? [X4] : (X3 != X4 & sK1 != sK2 & r2_projpl_1(sK0,sK2,X3,X4) & r2_projpl_1(sK0,sK1,X3,X4) & m1_subset_1(X4,u2_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) => (? [X4] : (sK3 != X4 & sK1 != sK2 & r2_projpl_1(sK0,sK2,sK3,X4) & r2_projpl_1(sK0,sK1,sK3,X4) & m1_subset_1(X4,u2_incsp_1(sK0))) & m1_subset_1(sK3,u2_incsp_1(sK0))) [choice axiom]
27. ? [X4] : (sK3 != X4 & sK1 != sK2 & r2_projpl_1(sK0,sK2,sK3,X4) & r2_projpl_1(sK0,sK1,sK3,X4) & m1_subset_1(X4,u2_incsp_1(sK0))) => (sK3 != sK4 & sK1 != sK2 & r2_projpl_1(sK0,sK2,sK3,sK4) & r2_projpl_1(sK0,sK1,sK3,sK4) & m1_subset_1(sK4,u2_incsp_1(sK0))) [choice axiom]
28. (((((sK3 != sK4 & sK1 != sK2 & r2_projpl_1(sK0,sK2,sK3,sK4) & r2_projpl_1(sK0,sK1,sK3,sK4) & m1_subset_1(sK4,u2_incsp_1(sK0))) & m1_subset_1(sK3,u2_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) | ~v1_projpl_1(sK0)) & (! [X5] : (! [X6] : (! [X7] : (! [X8] : (X7 = X8 | X5 = X6 | ~r2_projpl_1(sK0,X6,X7,X8) | ~r2_projpl_1(sK0,X5,X7,X8) | ~m1_subset_1(X8,u2_incsp_1(sK0))) | ~m1_subset_1(X7,u2_incsp_1(sK0))) | ~m1_subset_1(X6,u1_incsp_1(sK0))) | ~m1_subset_1(X5,u1_incsp_1(sK0))) | v1_projpl_1(sK0)) & l1_incsp_1(sK0) [skolemisation 22,27,26,25,24,23]
29. ! [X0] : (((v1_projpl_1(X0) | ? [X1] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & X1 != X2 & r1_incsp_1(X0,X2,X4) & r1_incsp_1(X0,X1,X4) & r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,X1,X3) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0)))) & (! [X1] : (! [X2] : (! [X3] : (! [X4] : (X3 = X4 | X1 = X2 | ~r1_incsp_1(X0,X2,X4) | ~r1_incsp_1(X0,X1,X4) | ~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3) | ~m1_subset_1(X4,u2_incsp_1(X0))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~v1_projpl_1(X0))) | ~l1_incsp_1(X0)) [nnf transformation 18]
30. ! [X0] : (((v1_projpl_1(X0) | ? [X1] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & X1 != X2 & r1_incsp_1(X0,X2,X4) & r1_incsp_1(X0,X1,X4) & r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,X1,X3) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0)))) & (! [X5] : (! [X6] : (! [X7] : (! [X8] : (X7 = X8 | X5 = X6 | ~r1_incsp_1(X0,X6,X8) | ~r1_incsp_1(X0,X5,X8) | ~r1_incsp_1(X0,X6,X7) | ~r1_incsp_1(X0,X5,X7) | ~m1_subset_1(X8,u2_incsp_1(X0))) | ~m1_subset_1(X7,u2_incsp_1(X0))) | ~m1_subset_1(X6,u1_incsp_1(X0))) | ~m1_subset_1(X5,u1_incsp_1(X0))) | ~v1_projpl_1(X0))) | ~l1_incsp_1(X0)) [rectify 29]
31. ! [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & X1 != X2 & r1_incsp_1(X0,X2,X4) & r1_incsp_1(X0,X1,X4) & r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,X1,X3) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) => (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & sK5(X0) != X2 & r1_incsp_1(X0,X2,X4) & r1_incsp_1(X0,sK5(X0),X4) & r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,sK5(X0),X3) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(sK5(X0),u1_incsp_1(X0)))) [choice axiom]
32. ! [X0] : (? [X2] : (? [X3] : (? [X4] : (X3 != X4 & sK5(X0) != X2 & r1_incsp_1(X0,X2,X4) & r1_incsp_1(X0,sK5(X0),X4) & r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,sK5(X0),X3) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) => (? [X3] : (? [X4] : (X3 != X4 & sK5(X0) != sK6(X0) & r1_incsp_1(X0,sK6(X0),X4) & r1_incsp_1(X0,sK5(X0),X4) & r1_incsp_1(X0,sK6(X0),X3) & r1_incsp_1(X0,sK5(X0),X3) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(sK6(X0),u1_incsp_1(X0)))) [choice axiom]
33. ! [X0] : (? [X3] : (? [X4] : (X3 != X4 & sK5(X0) != sK6(X0) & r1_incsp_1(X0,sK6(X0),X4) & r1_incsp_1(X0,sK5(X0),X4) & r1_incsp_1(X0,sK6(X0),X3) & r1_incsp_1(X0,sK5(X0),X3) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) => (? [X4] : (sK7(X0) != X4 & sK5(X0) != sK6(X0) & r1_incsp_1(X0,sK6(X0),X4) & r1_incsp_1(X0,sK5(X0),X4) & r1_incsp_1(X0,sK6(X0),sK7(X0)) & r1_incsp_1(X0,sK5(X0),sK7(X0)) & m1_subset_1(X4,u2_incsp_1(X0))) & m1_subset_1(sK7(X0),u2_incsp_1(X0)))) [choice axiom]
34. ! [X0] : (? [X4] : (sK7(X0) != X4 & sK5(X0) != sK6(X0) & r1_incsp_1(X0,sK6(X0),X4) & r1_incsp_1(X0,sK5(X0),X4) & r1_incsp_1(X0,sK6(X0),sK7(X0)) & r1_incsp_1(X0,sK5(X0),sK7(X0)) & m1_subset_1(X4,u2_incsp_1(X0))) => (sK7(X0) != sK8(X0) & sK5(X0) != sK6(X0) & r1_incsp_1(X0,sK6(X0),sK8(X0)) & r1_incsp_1(X0,sK5(X0),sK8(X0)) & r1_incsp_1(X0,sK6(X0),sK7(X0)) & r1_incsp_1(X0,sK5(X0),sK7(X0)) & m1_subset_1(sK8(X0),u2_incsp_1(X0)))) [choice axiom]
35. ! [X0] : (((v1_projpl_1(X0) | ((((sK7(X0) != sK8(X0) & sK5(X0) != sK6(X0) & r1_incsp_1(X0,sK6(X0),sK8(X0)) & r1_incsp_1(X0,sK5(X0),sK8(X0)) & r1_incsp_1(X0,sK6(X0),sK7(X0)) & r1_incsp_1(X0,sK5(X0),sK7(X0)) & m1_subset_1(sK8(X0),u2_incsp_1(X0))) & m1_subset_1(sK7(X0),u2_incsp_1(X0))) & m1_subset_1(sK6(X0),u1_incsp_1(X0))) & m1_subset_1(sK5(X0),u1_incsp_1(X0)))) & (! [X5] : (! [X6] : (! [X7] : (! [X8] : (X7 = X8 | X5 = X6 | ~r1_incsp_1(X0,X6,X8) | ~r1_incsp_1(X0,X5,X8) | ~r1_incsp_1(X0,X6,X7) | ~r1_incsp_1(X0,X5,X7) | ~m1_subset_1(X8,u2_incsp_1(X0))) | ~m1_subset_1(X7,u2_incsp_1(X0))) | ~m1_subset_1(X6,u1_incsp_1(X0))) | ~m1_subset_1(X5,u1_incsp_1(X0))) | ~v1_projpl_1(X0))) | ~l1_incsp_1(X0)) [skolemisation 30,34,33,32,31]
36. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r2_projpl_1(X0,X1,X2,X3) | (~r1_incsp_1(X0,X1,X3) | ~r1_incsp_1(X0,X1,X2))) & ((r1_incsp_1(X0,X1,X3) & r1_incsp_1(X0,X1,X2)) | ~r2_projpl_1(X0,X1,X2,X3))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u2_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [nnf transformation 19]
37. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r2_projpl_1(X0,X1,X2,X3) | ~r1_incsp_1(X0,X1,X3) | ~r1_incsp_1(X0,X1,X2)) & ((r1_incsp_1(X0,X1,X3) & r1_incsp_1(X0,X1,X2)) | ~r2_projpl_1(X0,X1,X2,X3))) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u2_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [flattening 36]
46. l1_incsp_1(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
47. X7 = X8 | X5 = X6 | ~r2_projpl_1(sK0,X6,X7,X8) | ~r2_projpl_1(sK0,X5,X7,X8) | ~m1_subset_1(X8,u2_incsp_1(sK0)) | ~m1_subset_1(X7,u2_incsp_1(sK0)) | ~m1_subset_1(X6,u1_incsp_1(sK0)) | ~m1_subset_1(X5,u1_incsp_1(sK0)) | v1_projpl_1(sK0) [cnf transformation 28] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
48. m1_subset_1(sK1,u1_incsp_1(sK0)) | ~v1_projpl_1(sK0) [cnf transformation 28] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
49. m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v1_projpl_1(sK0) [cnf transformation 28] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
50. m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v1_projpl_1(sK0) [cnf transformation 28] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
51. m1_subset_1(sK4,u2_incsp_1(sK0)) | ~v1_projpl_1(sK0) [cnf transformation 28] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
52. r2_projpl_1(sK0,sK1,sK3,sK4) | ~v1_projpl_1(sK0) [cnf transformation 28] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
53. r2_projpl_1(sK0,sK2,sK3,sK4) | ~v1_projpl_1(sK0) [cnf transformation 28] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
54. sK1 != sK2 | ~v1_projpl_1(sK0) [cnf transformation 28] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
55. sK3 != sK4 | ~v1_projpl_1(sK0) [cnf transformation 28] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
58. ~r1_incsp_1(X0,X6,X8) | X5 = X6 | X7 = X8 | ~r1_incsp_1(X0,X5,X8) | ~r1_incsp_1(X0,X6,X7) | ~r1_incsp_1(X0,X5,X7) | ~m1_subset_1(X8,u2_incsp_1(X0)) | ~m1_subset_1(X7,u2_incsp_1(X0)) | ~m1_subset_1(X6,u1_incsp_1(X0)) | ~m1_subset_1(X5,u1_incsp_1(X0)) | ~v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:42,wCS:16800,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
59. m1_subset_1(sK5(X0),u1_incsp_1(X0)) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
60. m1_subset_1(sK6(X0),u1_incsp_1(X0)) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
61. m1_subset_1(sK7(X0),u2_incsp_1(X0)) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
62. m1_subset_1(sK8(X0),u2_incsp_1(X0)) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
63. r1_incsp_1(X0,sK5(X0),sK7(X0)) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
64. r1_incsp_1(X0,sK6(X0),sK7(X0)) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
65. r1_incsp_1(X0,sK5(X0),sK8(X0)) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
66. r1_incsp_1(X0,sK6(X0),sK8(X0)) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
67. sK5(X0) != sK6(X0) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
68. sK7(X0) != sK8(X0) | v1_projpl_1(X0) | ~l1_incsp_1(X0) [cnf transformation 35] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
69. ~r2_projpl_1(X0,X1,X2,X3) | r1_incsp_1(X0,X1,X2) | ~m1_subset_1(X3,u2_incsp_1(X0)) | ~m1_subset_1(X2,u2_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) [cnf transformation 37] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. ~r2_projpl_1(X0,X1,X2,X3) | r1_incsp_1(X0,X1,X3) | ~m1_subset_1(X3,u2_incsp_1(X0)) | ~m1_subset_1(X2,u2_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) [cnf transformation 37] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. r2_projpl_1(X0,X1,X2,X3) | ~r1_incsp_1(X0,X1,X3) | ~r1_incsp_1(X0,X1,X2) | ~m1_subset_1(X3,u2_incsp_1(X0)) | ~m1_subset_1(X2,u2_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) [cnf transformation 37] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
77. 1 <=> v1_projpl_1(sK0) [avatar definition]
78. v1_projpl_1(sK0) <- (1) [avatar component clause 77] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
79. ~v1_projpl_1(sK0) <- (~1) [avatar component clause 77] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
81. 2 <=> sK3 = sK4 [avatar definition]
83. sK3 != sK4 <- (~2) [avatar component clause 81] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
84. ~1 | ~2 [avatar split clause 55,81,77]
86. 3 <=> sK1 = sK2 [avatar definition]
88. sK1 != sK2 <- (~3) [avatar component clause 86] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
89. ~1 | ~3 [avatar split clause 54,86,77]
91. 4 <=> r2_projpl_1(sK0,sK2,sK3,sK4) [avatar definition]
93. r2_projpl_1(sK0,sK2,sK3,sK4) <- (4) [avatar component clause 91] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
94. ~1 | 4 [avatar split clause 53,91,77]
96. 5 <=> r2_projpl_1(sK0,sK1,sK3,sK4) [avatar definition]
98. r2_projpl_1(sK0,sK1,sK3,sK4) <- (5) [avatar component clause 96] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
99. ~1 | 5 [avatar split clause 52,96,77]
101. 6 <=> m1_subset_1(sK4,u2_incsp_1(sK0)) [avatar definition]
103. m1_subset_1(sK4,u2_incsp_1(sK0)) <- (6) [avatar component clause 101] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
104. ~1 | 6 [avatar split clause 51,101,77]
106. 7 <=> m1_subset_1(sK3,u2_incsp_1(sK0)) [avatar definition]
108. m1_subset_1(sK3,u2_incsp_1(sK0)) <- (7) [avatar component clause 106] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
109. ~1 | 7 [avatar split clause 50,106,77]
111. 8 <=> m1_subset_1(sK2,u1_incsp_1(sK0)) [avatar definition]
113. m1_subset_1(sK2,u1_incsp_1(sK0)) <- (8) [avatar component clause 111] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
114. ~1 | 8 [avatar split clause 49,111,77]
116. 9 <=> m1_subset_1(sK1,u1_incsp_1(sK0)) [avatar definition]
118. m1_subset_1(sK1,u1_incsp_1(sK0)) <- (9) [avatar component clause 116] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
119. ~1 | 9 [avatar split clause 48,116,77]
121. 10 <=> ! [X5,X7,X8,X6] : (X7 = X8 | ~m1_subset_1(X5,u1_incsp_1(sK0)) | ~m1_subset_1(X6,u1_incsp_1(sK0)) | ~m1_subset_1(X7,u2_incsp_1(sK0)) | ~m1_subset_1(X8,u2_incsp_1(sK0)) | ~r2_projpl_1(sK0,X5,X7,X8) | ~r2_projpl_1(sK0,X6,X7,X8) | X5 = X6) [avatar definition]
122. ~r2_projpl_1(sK0,X6,X7,X8) | ~m1_subset_1(X5,u1_incsp_1(sK0)) | ~m1_subset_1(X6,u1_incsp_1(sK0)) | ~m1_subset_1(X7,u2_incsp_1(sK0)) | ~m1_subset_1(X8,u2_incsp_1(sK0)) | ~r2_projpl_1(sK0,X5,X7,X8) | X7 = X8 | X5 = X6 <- (10) [avatar component clause 121] {a:0,w:32,wCS:3200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
123. 1 | 10 [avatar split clause 47,121,77]
125. 11 <=> l1_incsp_1(sK0) [avatar definition]
127. l1_incsp_1(sK0) <- (11) [avatar component clause 125] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
128. 11 [avatar split clause 46,125]
144. ~r1_incsp_1(sK0,X0,X1) | ~r1_incsp_1(sK0,X0,X2) | ~m1_subset_1(X1,u2_incsp_1(sK0)) | ~m1_subset_1(X2,u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~m1_subset_1(X3,u1_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~m1_subset_1(X2,u2_incsp_1(sK0)) | ~m1_subset_1(X1,u2_incsp_1(sK0)) | ~r2_projpl_1(sK0,X3,X2,X1) | X1 = X2 | X0 = X3 <- (10) [resolution 71,122] {a:1,w:49,wCS:4900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
149. ~r1_incsp_1(sK0,X0,X1) | ~r1_incsp_1(sK0,X0,X2) | ~m1_subset_1(X1,u2_incsp_1(sK0)) | ~m1_subset_1(X2,u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~m1_subset_1(X3,u1_incsp_1(sK0)) | ~r2_projpl_1(sK0,X3,X2,X1) | X1 = X2 | X0 = X3 <- (10) [duplicate literal removal 144] {a:1,w:37,wCS:3700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
150. ~r2_projpl_1(sK0,X3,X2,X1) | ~r1_incsp_1(sK0,X0,X2) | ~m1_subset_1(X1,u2_incsp_1(sK0)) | ~m1_subset_1(X2,u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~m1_subset_1(X3,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X0,X1) | X1 = X2 | X0 = X3 <- (10, 11) [subsumption resolution 149,127] {a:1,w:35,wCS:3500,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
159. ~r1_incsp_1(sK0,X0,X1) | ~m1_subset_1(X2,u2_incsp_1(sK0)) | ~m1_subset_1(X1,u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~m1_subset_1(X3,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X0,X2) | X1 = X2 | X0 = X3 | ~r1_incsp_1(sK0,X3,X2) | ~r1_incsp_1(sK0,X3,X1) | ~m1_subset_1(X2,u2_incsp_1(sK0)) | ~m1_subset_1(X1,u2_incsp_1(sK0)) | ~m1_subset_1(X3,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (10, 11) [resolution 150,71] {a:2,w:52,wCS:5200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
160. ~r1_incsp_1(sK0,X0,X1) | ~m1_subset_1(X2,u2_incsp_1(sK0)) | ~m1_subset_1(X1,u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~m1_subset_1(X3,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X0,X2) | X1 = X2 | X0 = X3 | ~r1_incsp_1(sK0,X3,X2) | ~r1_incsp_1(sK0,X3,X1) | ~l1_incsp_1(sK0) <- (10, 11) [duplicate literal removal 159] {a:2,w:40,wCS:4000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
161. ~r1_incsp_1(sK0,X3,X2) | ~m1_subset_1(X2,u2_incsp_1(sK0)) | ~m1_subset_1(X1,u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~m1_subset_1(X3,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X0,X2) | X1 = X2 | X0 = X3 | ~r1_incsp_1(sK0,X0,X1) | ~r1_incsp_1(sK0,X3,X1) <- (10, 11) [subsumption resolution 160,127] {a:2,w:38,wCS:3800,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
164. ~m1_subset_1(sK8(sK0),u2_incsp_1(sK0)) | ~m1_subset_1(X4,u2_incsp_1(sK0)) | ~m1_subset_1(X5,u1_incsp_1(sK0)) | ~m1_subset_1(sK5(sK0),u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X5,sK8(sK0)) | sK8(sK0) = X4 | sK5(sK0) = X5 | ~r1_incsp_1(sK0,X5,X4) | ~r1_incsp_1(sK0,sK5(sK0),X4) | v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (10, 11) [resolution 161,65] {a:3,w:44,wCS:4400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
169. 15 <=> m1_subset_1(sK5(sK0),u1_incsp_1(sK0)) [avatar definition]
171. ~m1_subset_1(sK5(sK0),u1_incsp_1(sK0)) <- (~15) [avatar component clause 169] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
176. 17 <=> m1_subset_1(sK7(sK0),u2_incsp_1(sK0)) [avatar definition]
178. ~m1_subset_1(sK7(sK0),u2_incsp_1(sK0)) <- (~17) [avatar component clause 176] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
183. 18 <=> m1_subset_1(sK6(sK0),u1_incsp_1(sK0)) [avatar definition]
185. ~m1_subset_1(sK6(sK0),u1_incsp_1(sK0)) <- (~18) [avatar component clause 183] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
193. 20 <=> ! [X5,X4] : (~m1_subset_1(X4,u2_incsp_1(sK0)) | ~r1_incsp_1(sK0,sK5(sK0),X4) | sK8(sK0) = X4 | ~m1_subset_1(X5,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X5,X4) | sK5(sK0) = X5 | ~r1_incsp_1(sK0,X5,sK8(sK0))) [avatar definition]
194. ~r1_incsp_1(sK0,sK5(sK0),X4) | ~m1_subset_1(X4,u2_incsp_1(sK0)) | sK8(sK0) = X4 | ~m1_subset_1(X5,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X5,X4) | sK5(sK0) = X5 | ~r1_incsp_1(sK0,X5,sK8(sK0)) <- (20) [avatar component clause 193] {a:3,w:30,wCS:3000,nSel:1,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
196. 21 <=> m1_subset_1(sK8(sK0),u2_incsp_1(sK0)) [avatar definition]
198. ~m1_subset_1(sK8(sK0),u2_incsp_1(sK0)) <- (~21) [avatar component clause 196] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
207. v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (~15) [resolution 171,59] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
212. v1_projpl_1(sK0) <- (11, ~15) [subsumption resolution 207,127] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
213. 1 | ~11 | 15 [avatar split clause 212,169,125,77]
220. r1_incsp_1(sK0,sK2,sK4) | ~m1_subset_1(sK4,u2_incsp_1(sK0)) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4) [resolution 93,70] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
221. r1_incsp_1(sK0,sK2,sK3) | ~m1_subset_1(sK4,u2_incsp_1(sK0)) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4) [resolution 93,69] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
230. r1_incsp_1(sK0,sK2,sK4) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 6) [subsumption resolution 220,103] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
235. 23 <=> r1_incsp_1(sK0,sK2,sK4) [avatar definition]
237. r1_incsp_1(sK0,sK2,sK4) <- (23) [avatar component clause 235] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
239. r1_incsp_1(sK0,sK2,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 6) [subsumption resolution 221,103] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
244. 24 <=> r1_incsp_1(sK0,sK2,sK3) [avatar definition]
246. r1_incsp_1(sK0,sK2,sK3) <- (24) [avatar component clause 244] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
248. r1_incsp_1(sK0,sK2,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 6, 8) [subsumption resolution 239,113] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
249. r1_incsp_1(sK0,sK2,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) <- (4, 6, 8, 11) [subsumption resolution 248,127] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
250. ~7 | 24 | ~4 | ~6 | ~8 | ~11 [avatar split clause 249,125,111,101,91,244,106]
251. r1_incsp_1(sK0,sK2,sK4) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 6, 8) [subsumption resolution 230,113] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
252. r1_incsp_1(sK0,sK2,sK4) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) <- (4, 6, 8, 11) [subsumption resolution 251,127] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
253. ~7 | 23 | ~4 | ~6 | ~8 | ~11 [avatar split clause 252,125,111,101,91,235,106]
268. ~m1_subset_1(sK8(sK0),u2_incsp_1(sK0)) | ~m1_subset_1(X4,u2_incsp_1(sK0)) | ~m1_subset_1(X5,u1_incsp_1(sK0)) | ~m1_subset_1(sK5(sK0),u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X5,sK8(sK0)) | sK8(sK0) = X4 | sK5(sK0) = X5 | ~r1_incsp_1(sK0,X5,X4) | ~r1_incsp_1(sK0,sK5(sK0),X4) | v1_projpl_1(sK0) <- (10, 11) [subsumption resolution 164,127] {a:3,w:42,wCS:4200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
269. 1 | ~15 | 20 | ~21 | ~10 | ~11 [avatar split clause 268,125,121,196,193,169,77]
272. r1_incsp_1(sK0,sK1,sK4) | ~m1_subset_1(sK4,u2_incsp_1(sK0)) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (5) [resolution 98,70] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
273. r1_incsp_1(sK0,sK1,sK3) | ~m1_subset_1(sK4,u2_incsp_1(sK0)) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (5) [resolution 98,69] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
275. v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (~17) [resolution 178,61] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
277. ~l1_incsp_1(sK0) <- (~1, ~17) [subsumption resolution 275,79] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
278. $false <- (~1, 11, ~17) [subsumption resolution 277,127] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
279. 1 | ~11 | 17 [avatar contradiction clause 278]
281. v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (~18) [resolution 185,60] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
283. ~l1_incsp_1(sK0) <- (~1, ~18) [subsumption resolution 281,79] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
284. $false <- (~1, 11, ~18) [subsumption resolution 283,127] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
285. 1 | ~11 | 18 [avatar contradiction clause 284]
287. v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (~21) [resolution 198,62] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
289. ~l1_incsp_1(sK0) <- (~1, ~21) [subsumption resolution 287,79] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
290. $false <- (~1, 11, ~21) [subsumption resolution 289,127] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
291. 1 | ~11 | 21 [avatar contradiction clause 290]
298. 27 <=> ! [X1] : (~m1_subset_1(X1,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X1,sK7(sK0)) | sK5(sK0) = X1 | ~r1_incsp_1(sK0,X1,sK8(sK0))) [avatar definition]
299. ~r1_incsp_1(sK0,X1,sK8(sK0)) | ~r1_incsp_1(sK0,X1,sK7(sK0)) | sK5(sK0) = X1 | ~m1_subset_1(X1,u1_incsp_1(sK0)) <- (27) [avatar component clause 298] {a:4,w:18,wCS:1800,nSel:1,goal:1,sine:255,thAx:0,allAx:18,thDist:-18}
301. 28 <=> sK7(sK0) = sK8(sK0) [avatar definition]
303. sK7(sK0) = sK8(sK0) <- (28) [avatar component clause 301] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:18,thDist:-18}
317. sK7(sK0) != sK7(sK0) | v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (28) [superposition 68,303] {a:5,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
319. v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (28) [trivial inequality removal 317] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
321. ~l1_incsp_1(sK0) <- (~1, 28) [subsumption resolution 319,79] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
322. $false <- (~1, 11, 28) [subsumption resolution 321,127] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
323. 1 | ~11 | ~28 [avatar contradiction clause 322]
324. ~m1_subset_1(sK7(sK0),u2_incsp_1(sK0)) | sK7(sK0) = sK8(sK0) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X0,sK7(sK0)) | sK5(sK0) = X0 | ~r1_incsp_1(sK0,X0,sK8(sK0)) | v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (20) [resolution 194,63] {a:4,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
327. ~m1_subset_1(sK7(sK0),u2_incsp_1(sK0)) | sK7(sK0) = sK8(sK0) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X0,sK7(sK0)) | sK5(sK0) = X0 | ~r1_incsp_1(sK0,X0,sK8(sK0)) | ~l1_incsp_1(sK0) <- (~1, 20) [subsumption resolution 324,79] {a:4,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
328. ~m1_subset_1(sK7(sK0),u2_incsp_1(sK0)) | sK7(sK0) = sK8(sK0) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~r1_incsp_1(sK0,X0,sK7(sK0)) | sK5(sK0) = X0 | ~r1_incsp_1(sK0,X0,sK8(sK0)) <- (~1, 11, 20) [subsumption resolution 327,127] {a:4,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
329. 27 | 28 | ~17 | 1 | ~11 | ~20 [avatar split clause 328,193,125,77,176,301,298]
337. ~r1_incsp_1(sK0,sK6(sK0),sK7(sK0)) | sK5(sK0) = sK6(sK0) | ~m1_subset_1(sK6(sK0),u1_incsp_1(sK0)) | v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (27) [resolution 299,66] {a:5,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
341. 30 <=> sK5(sK0) = sK6(sK0) [avatar definition]
343. sK5(sK0) = sK6(sK0) <- (30) [avatar component clause 341] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:23,thDist:-23}
345. 31 <=> r1_incsp_1(sK0,sK6(sK0),sK7(sK0)) [avatar definition]
347. ~r1_incsp_1(sK0,sK6(sK0),sK7(sK0)) <- (~31) [avatar component clause 345] {a:5,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:23,thDist:-23}
360. v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (~31) [resolution 347,64] {a:6,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:28,thDist:-28}
362. ~l1_incsp_1(sK0) <- (~1, ~31) [subsumption resolution 360,79] {a:6,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:28,thDist:-28}
363. $false <- (~1, 11, ~31) [subsumption resolution 362,127] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:28,thDist:-28}
364. 1 | ~11 | 31 [avatar contradiction clause 363]
371. sK5(sK0) != sK5(sK0) | v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (30) [superposition 67,343] {a:6,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:28,thDist:-28}
373. v1_projpl_1(sK0) | ~l1_incsp_1(sK0) <- (30) [trivial inequality removal 371] {a:6,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:28,thDist:-28}
377. ~l1_incsp_1(sK0) <- (~1, 30) [subsumption resolution 373,79] {a:6,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:28,thDist:-28}
378. $false <- (~1, 11, 30) [subsumption resolution 377,127] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:28,thDist:-28}
379. 1 | ~11 | ~30 [avatar contradiction clause 378]
382. ~r1_incsp_1(sK0,sK6(sK0),sK7(sK0)) | sK5(sK0) = sK6(sK0) | ~m1_subset_1(sK6(sK0),u1_incsp_1(sK0)) | v1_projpl_1(sK0) <- (11, 27) [subsumption resolution 337,127] {a:5,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
383. 1 | ~18 | 30 | ~31 | ~11 | ~27 [avatar split clause 382,298,125,345,341,183,77]
388. r1_incsp_1(sK0,sK1,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (5, 6) [subsumption resolution 273,103] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
389. r1_incsp_1(sK0,sK1,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (5, 6, 9) [subsumption resolution 388,118] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
390. r1_incsp_1(sK0,sK1,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) <- (5, 6, 9, 11) [subsumption resolution 389,127] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
392. 33 <=> r1_incsp_1(sK0,sK1,sK3) [avatar definition]
394. r1_incsp_1(sK0,sK1,sK3) <- (33) [avatar component clause 392] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
395. ~7 | 33 | ~5 | ~6 | ~9 | ~11 [avatar split clause 390,125,116,101,96,392,106]
396. r1_incsp_1(sK0,sK1,sK4) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (5, 6) [subsumption resolution 272,103] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
397. r1_incsp_1(sK0,sK1,sK4) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (5, 6, 9) [subsumption resolution 396,118] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
398. r1_incsp_1(sK0,sK1,sK4) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) <- (5, 6, 9, 11) [subsumption resolution 397,127] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
400. 34 <=> r1_incsp_1(sK0,sK1,sK4) [avatar definition]
402. r1_incsp_1(sK0,sK1,sK4) <- (34) [avatar component clause 400] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
403. ~7 | 34 | ~5 | ~6 | ~9 | ~11 [avatar split clause 398,125,116,101,96,400,106]
440. $false <- (1, ~2, ~3, 6, 7, 8, 9, 11, 23, 24, 33, 34) [unit resulting resolution 127,78,88,83,394,108,246,113,237,118,103,402,58] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:81,thDist:-81}
445. ~1 | 2 | 3 | ~6 | ~7 | ~8 | ~9 | ~11 | ~23 | ~24 | ~33 | ~34 [avatar contradiction clause 440]
450. $false [avatar sat refutation 84,89,94,99,104,109,114,119,123,128,213,250,253,269,279,285,291,323,329,364,379,383,395,403,445] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1497,thDist:-1497}
% SZS output end Proof for projpl_1__t3_projpl_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 13
% 
% >>> Saturation
% Initial clauses: 30
% Generated clauses: 275
% Active clauses: 56
% Passive clauses: 69
% Final active clauses: 38
% Final passive clauses: 11
% Main loop iterations started: 56
% 
% >>> Simplifying Inferences
% Duplicate literals: 22
% Trivial inequalities: 2
% Fw subsumption resolutions: 98
% Fw demodulations: 2
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Equational tautologies: 6
% Forward subsumptions: 17
% 
% >>> Generating Inferences
% Binary resolution: 41
% Unit resulting resolution: 13
% Backward superposition: 8
% 
% >>> AVATAR
% Split clauses: 35
% Split components: 94
% Sat splitting refutations: 20
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 64
% SAT solver unit clauses: 4
% SAT solver binary clauses: 9
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% splitting component index usage: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

