i: [1,166,0,2,1,1,0,0] t25_relset_1
a: [3,270,0,2,1,166]
s: 270
i: [1,165,0,2,1,1,0,0] t25_relset_1
a: [3,275,0,2,1,165]
s: 275
i: [1,168,0,8,2,1,0,0] t25_relset_1
a: [3,260,0,3,1,168]
s: 260
i: [1,170,0,13,3,1,0,0] t25_relset_1
a: [3,247,0,5,1,170]
s: 247
i: [1,169,0,10,2,1,0,0] t25_relset_1
a: [3,255,0,5,1,169]
s: 255
i: [1,167,0,6,1,1,0,0] t25_relset_1
a: [3,265,0,6,1,167]
s: 265
i: [1,171,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,280,0,2,1,171]
s: 280
i: [1,172,0,2,1,0,0,255] fc1_xboole_0
a: [3,285,0,2,1,172]
s: 285
i: [1,237,0,2,1,0,0,1] rc2_xboole_0
a: [3,295,0,2,1,237]
s: 295
i: [1,238,0,2,1,0,0,1] rc1_xboole_0
a: [3,300,0,2,1,238]
s: 300
i: [1,240,0,2,1,0,0,2] rc1_relat_1
a: [3,305,0,2,1,240]
s: 305
i: [1,239,0,2,1,0,0,2] rc1_relat_1
a: [3,310,0,2,1,239]
s: 310
i: [1,241,0,2,1,0,0,255] rc2_relat_1
a: [3,315,0,2,1,241]
s: 315
i: [1,174,0,3,1,0,0,255] fc2_xboole_0
s: 174
i: [1,175,0,3,1,0,0,1] fc1_subset_1
s: 175
i: [1,214,0,3,1,0,0,1] rc2_subset_1
s: 214
i: [1,217,0,3,1,0,0,255] reflexivity_r1_tarski
s: 217
i: [1,173,0,3,1,0,0,255] d2_xboole_0
a: [3,290,0,3,1,173]
s: 290
r: 280
i: [1,183,0,4,2,0,0,2] cc1_relat_1
s: 183
i: [1,212,0,4,1,0,0,1] existence_m1_subset_1
s: 212
i: [1,216,0,4,1,0,0,255] rc3_subset_1
s: 216
d: [2,318,1,2,1,1,40,183,285]
a: [3,324,1,2,1,318]
s: 324
i: [1,218,0,4,1,0,0,255] fc3_xboole_0
s: 218
d: [2,320,1,2,1,1,40,183,300]
a: [3,329,1,2,1,320]
s: 329
i: [1,219,0,4,1,0,0,1] fc6_relat_1
s: 219
i: [1,233,0,4,1,0,0,255] rc3_relat_1
s: 233
d: [2,319,1,3,1,0,40,183,214]
s: 319
i: [1,179,0,5,2,0,0,1] rc1_subset_1
s: 179
i: [1,184,0,5,2,0,0,255] t6_boole
s: 184
i: [1,185,0,5,2,0,0,2] fc11_relat_1
s: 185
i: [1,206,0,5,2,0,0,1] t7_boole
s: 206
d: [2,335,1,2,1,1,40,206,255]
a: [3,339,1,2,1,335]
s: 339
i: [1,213,0,5,1,0,0,1] rc2_subset_1
s: 213
i: [1,215,0,5,1,0,0,255] rc3_subset_1
s: 215
i: [1,220,0,5,1,0,0,255] fc5_relat_1
s: 220
i: [1,234,0,5,1,0,0,255] rc3_relat_1
s: 234
i: [1,235,0,5,1,0,0,255] rc3_relat_1
s: 235
i: [1,181,0,6,2,0,0,255] rc4_subset_1
s: 181
d: [2,334,1,5,2,0,40,185,183]
s: 334
i: [1,198,0,6,2,0,0,1] t1_subset
s: 198
r: 247
r: 255
a: [3,246,0,5,1,170]
s: 246
r: 246
s: 247
s: 255
d: [2,341,1,5,1,1,40,198,255]
a: [3,345,1,5,1,341]
s: 345
a: [3,250,0,8,2,170]
s: 250
d: [2,362,1,0,0,3,52,260,255,250]
e: 362
d: [2,363,1,3,1,2,40,250,255]
d: [2,365,1,0,0,3,21,363,260]
e: 365
r: 247
r: 255
s: 246
d: [2,348,1,4,1,1,40,246,206]
a: [3,357,1,4,1,348]
s: 357
i: [1,224,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 224
d: [2,347,1,5,1,1,40,246,198]
a: [3,352,1,5,1,347]
s: 352
i: [1,178,0,7,2,0,0,1] rc1_subset_1
s: 178
i: [1,180,0,7,2,0,0,255] rc4_subset_1
s: 180
d: [2,371,1,5,1,1,40,224,246]
a: [3,375,1,5,1,371]
s: 375
i: [1,194,0,7,3,0,0,2] fc9_relat_1
s: 194
i: [1,204,0,7,2,0,0,255] t3_subset
s: 204
i: [1,205,0,7,3,0,0,2] t8_boole
s: 205
i: [1,221,0,7,1,0,0,255] commutativity_k2_tarski
s: 221
d: [2,331,1,7,3,0,44,184,184]
s: 331
i: [1,182,0,8,3,0,0,2] cc2_relat_1
s: 182
d: [2,381,1,6,2,1,40,182,265]
d: [2,387,1,2,1,1,21,381,219]
a: [3,391,1,2,1,387]
s: 391
i: [1,186,0,8,3,0,0,1] cc1_subset_1
s: 186
i: [1,199,0,8,3,0,0,1] t2_subset
s: 199
r: 339
a: [3,338,1,2,1,335]
s: 338
a: [3,254,0,5,1,169]
s: 254
d: [2,410,1,5,2,1,40,199,260]
d: [2,425,1,3,1,2,21,410,270]
a: [3,429,1,3,1,425]
s: 429
i: [1,225,0,8,3,0,0,1] fc10_subset_1
s: 225
i: [1,230,0,8,2,0,0,1] cc1_relset_1
s: 230
d: [2,438,2,3,1,1,40,429,224]
a: [3,444,2,3,1,438]
s: 444
i: [1,236,0,8,1,0,0,255] fc7_relat_1
s: 236
d: [2,394,1,6,2,1,40,186,265]
d: [2,400,1,4,1,2,21,394,339]
a: [3,404,1,4,1,400]
s: 404
i: [1,187,0,9,3,0,0,255] cc4_subset_1
s: 187
i: [1,188,0,9,4,0,0,255] cc7_relat_1
s: 188
d: [2,408,1,11,2,1,40,199,265]
d: [2,419,1,6,1,1,21,408,175]
a: [3,423,1,6,1,419]
s: 423
i: [1,191,0,9,4,0,0,2] cc8_relat_1
s: 191
d: [2,464,2,6,1,1,40,423,224]
a: [3,470,2,6,1,464]
s: 470
i: [1,208,0,9,3,0,0,1] t5_subset
s: 208
i: [1,231,0,9,2,0,0,255] cc2_relset_1
s: 231
d: [2,482,1,3,1,1,40,231,265]
a: [3,491,1,3,1,482]
s: 491
i: [1,232,0,9,2,0,0,255] cc2_relset_1
s: 232
d: [2,495,1,3,1,1,40,232,265]
a: [3,504,1,3,1,495]
s: 504
i: [1,242,0,9,2,0,0,2] d13_xtuple_0
s: 242
i: [1,197,0,10,1,0,0,255] d5_tarski
s: 197
d: [2,514,1,4,1,0,44,218,197]
s: 514
i: [1,207,0,10,3,0,0,1] t4_subset
s: 207
d: [2,516,1,8,2,1,40,207,265]
s: 516
i: [1,209,0,10,2,0,0,1] t87_zfmisc_1
s: 209
i: [1,210,0,10,2,0,0,1] t87_zfmisc_1
s: 210
d: [2,522,2,12,3,1,40,516,199]
d: [2,523,2,8,2,2,21,522,404]
s: 523
i: [1,222,0,10,3,0,0,1] cc4_relset_1
s: 222
d: [2,524,3,8,2,2,40,523,224]
s: 524
i: [1,223,0,10,3,0,0,1] cc3_relset_1
s: 223
i: [1,176,0,11,4,0,0,255] cc3_subset_1
s: 176
r: 524
r: 523
r: 404
s: 404
i: [1,177,0,11,4,0,0,255] cc2_subset_1
s: 177
d: [2,554,1,11,3,1,40,176,265]
d: [2,562,1,9,2,2,21,554,338]
d: [2,563,1,5,1,3,21,562,404]
a: [3,567,1,5,1,563]
s: 567
i: [1,227,0,11,3,0,0,1] redefinition_k2_relset_1
s: 227
r: 246
r: 250
r: 254
r: 338
r: 352
s: 247
s: 339
s: 255
r: 247
r: 255
r: 339
s: 338
s: 246
s: 254
d: [2,583,2,8,3,1,44,357,227]
d: [2,599,2,5,2,2,21,583,504]
d: [2,600,2,3,1,3,21,599,391]
a: [3,604,2,3,1,600]
s: 604
d: [2,636,3,0,0,2,52,338,604,185]
e: 636
d: [2,637,3,2,1,1,40,604,185]
d: [2,639,3,0,0,2,21,637,338]
e: 639
r: 338
r: 345
s: 250
s: 339
i: [1,228,0,11,3,0,0,1] dt_k2_relset_1
s: 228
i: [1,243,0,11,2,0,0,2] d13_xtuple_0
s: 243
d: [2,584,1,9,3,1,44,246,227]
d: [2,606,1,6,2,2,21,584,504]
d: [2,607,1,4,1,3,21,606,391]
a: [3,611,1,4,1,607]
s: 611
i: [1,226,0,12,4,0,0,255] cc5_relat_1
s: 226
d: [2,581,2,9,3,1,44,375,227]
d: [2,585,2,6,2,2,21,581,504]
d: [2,586,2,4,1,3,21,585,391]
a: [3,590,2,4,1,586]
s: 590
i: [1,229,0,12,4,0,0,2] cc6_relat_1
s: 229
i: [1,211,0,13,3,0,0,1] t87_zfmisc_1
s: 211
d: [2,582,2,9,3,1,44,352,227]
d: [2,592,2,6,2,2,21,582,504]
d: [2,593,2,4,1,3,21,592,391]
a: [3,597,2,4,1,593]
s: 597
i: [1,203,0,16,3,0,0,2] d13_xtuple_0
s: 203
a: [3,344,1,5,1,341]
s: 344
i: [1,202,0,18,3,0,0,2] d13_xtuple_0
s: 202
s: 352
d: [2,660,1,9,2,1,40,243,250]
a: [3,667,1,5,1,660]
s: 667
d: [2,676,1,12,3,1,40,226,265]
d: [2,682,1,8,2,1,21,676,219]
s: 682
d: [2,538,4,10,2,2,40,524,523]
d: [2,539,4,5,1,2,18,538]
a: [3,543,4,5,1,539]
s: 543
d: [2,685,1,12,3,1,40,229,265]
d: [2,691,1,8,2,1,21,685,219]
s: 691
d: [2,386,1,5,2,0,40,182,215]
s: 386
a: [3,574,2,8,2,522]
s: 574
d: [2,399,1,5,2,0,40,186,215]
s: 399
d: [2,711,3,8,2,1,40,574,224]
s: 711
d: [2,414,1,8,2,0,40,199,213]
d: [2,433,1,5,1,0,21,414,175]
s: 433
d: [2,415,1,8,2,0,40,199,215]
d: [2,434,1,5,1,0,21,415,175]
s: 434
d: [2,714,3,8,2,1,40,574,210]
s: 714
d: [2,452,1,5,1,0,40,230,215]
s: 452
d: [2,727,4,7,2,1,40,714,243]
s: 727
d: [2,559,1,9,3,0,40,176,215]
d: [2,570,1,5,2,0,21,559,216]
s: 570
d: [2,384,1,6,2,0,40,182,212]
s: 384
d: [2,731,5,3,1,2,40,727,611]
a: [3,735,5,3,1,731]
s: 735
d: [2,397,1,6,2,0,40,186,212]
s: 397
d: [2,738,6,3,1,1,40,735,224]
a: [3,744,6,3,1,738]
s: 744
d: [2,413,1,6,2,0,40,199,212]
s: 413
d: [2,450,1,6,1,0,40,230,212]
s: 450
d: [2,739,6,3,1,1,40,735,198]
a: [3,749,6,3,1,739]
s: 749
d: [2,460,1,6,2,0,40,187,213]
s: 460
d: [2,755,5,8,2,1,40,413,727]
d: [2,756,5,5,1,2,21,755,604]
a: [3,760,5,5,1,756]
s: 760
d: [2,462,1,10,3,0,40,188,234]
d: [2,463,1,6,2,0,21,462,233]
s: 463
d: [2,472,1,10,3,0,40,191,235]
d: [2,473,1,6,2,0,21,472,233]
s: 473
d: [2,763,6,5,1,1,40,760,224]
a: [3,769,6,5,1,763]
s: 769
d: [2,479,1,6,2,0,40,208,213]
s: 479
d: [2,764,6,5,1,1,40,760,198]
a: [3,774,6,5,1,764]
s: 774
d: [2,480,1,6,2,0,40,208,215]
s: 480
d: [2,486,1,6,1,0,40,231,213]
s: 486
d: [2,715,3,8,2,1,40,574,209]
s: 715
d: [2,487,1,6,1,0,40,231,215]
s: 487
d: [2,786,4,9,2,1,40,715,243]
s: 786
d: [2,499,1,6,1,0,40,232,213]
s: 499
d: [2,500,1,6,1,0,40,232,215]
s: 500
d: [2,790,5,9,2,1,40,786,224]
s: 790
d: [2,532,1,11,3,0,40,222,178]
d: [2,537,1,6,2,0,21,532,179]
s: 537
d: [2,791,5,9,2,1,40,786,198]
s: 791
d: [2,799,6,0,0,3,52,611,667,791]
e: 799
d: [2,800,6,4,1,2,40,791,667]
d: [2,803,6,0,0,3,21,800,611]
e: 803
d: [2,368,1,4,1,1,40,246,206]
d: [2,613,1,9,3,1,44,247,227]
d: [2,614,1,6,2,2,21,613,504]
d: [2,615,1,4,1,3,21,614,391]
d: [2,617,1,4,1,1,40,255,242]
d: [2,631,1,9,3,1,44,246,227]
d: [2,633,1,6,2,2,21,631,504]
d: [2,634,1,4,1,3,21,633,391]
f: 803,660,637,634,617,615,600,522,495,394,387,368,167,169,170
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relset_1__t25_relset_1
% SZS output start Proof for relset_1__t25_relset_1
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X0))) => ! [X3] : (r2_hidden(X3,k2_relset_1(X0,X2)) <=> ? [X4] : (r2_hidden(k4_tarski(X4,X3),X2) & m1_subset_1(X4,X1)))))) [input t25_relset_1]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X0))) => ! [X3] : (r2_hidden(X3,k2_relset_1(X0,X2)) <=> ? [X4] : (r2_hidden(k4_tarski(X4,X3),X2) & m1_subset_1(X4,X1)))))) [negated conjecture 1]
6. ! [X0] : (v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_xboole_0(X1))) [input cc1_subset_1]
7. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
8. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [input cc2_relset_1]
21. ! [X0] : ! [X1] : (k10_xtuple_0(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : r2_hidden(k4_tarski(X3,X2),X0))) [input d13_xtuple_0]
36. ! [X0] : (v1_xboole_0(X0) => v1_xboole_0(k10_xtuple_0(X0))) [input fc11_relat_1]
42. ! [X0,X1] : v1_relat_1(k2_zfmisc_1(X0,X1)) [input fc6_relat_1]
54. ! [X0,X1] : ((v5_relat_1(X1,X0) & v1_relat_1(X1)) => k2_relset_1(X0,X1) = k10_xtuple_0(X1)) [input redefinition_k2_relset_1]
56. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
57. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
59. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
62. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
63. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [input t87_zfmisc_1]
66. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 56]
67. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 57]
68. ! [X0,X1] : (k10_xtuple_0(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : r2_hidden(k4_tarski(X3,X2),X0))) [flattening 21]
71. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 62]
72. ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 59]
74. ! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [flattening 63]
77. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [flattening 8]
82. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (r2_hidden(X3,k2_relset_1(X0,X2)) <~> ? [X4] : (r2_hidden(k4_tarski(X4,X3),X2) & m1_subset_1(X4,X1))) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X0)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [ennf transformation 2]
89. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 7]
92. ! [X0] : (v1_xboole_0(k10_xtuple_0(X0)) | ~v1_xboole_0(X0)) [ennf transformation 36]
93. ! [X0] : (! [X1] : (v1_xboole_0(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_xboole_0(X0)) [ennf transformation 6]
105. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 66]
106. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 67]
107. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 106]
110. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 71]
111. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))) [ennf transformation 72]
112. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [flattening 111]
121. ! [X0,X1] : (k2_relset_1(X0,X1) = k10_xtuple_0(X1) | (~v5_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 54]
122. ! [X0,X1] : (k2_relset_1(X0,X1) = k10_xtuple_0(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 121]
128. ! [X0,X1,X2] : ((v5_relat_1(X2,X1) & v4_relat_1(X2,X0)) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 77]
129. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X4,X3),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k2_relset_1(X0,X2))) & (? [X4] : (r2_hidden(k4_tarski(X4,X3),X2) & m1_subset_1(X4,X1)) | r2_hidden(X3,k2_relset_1(X0,X2)))) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X0)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [nnf transformation 82]
130. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X4,X3),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k2_relset_1(X0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X5,X3),X2) & m1_subset_1(X5,X1)) | r2_hidden(X3,k2_relset_1(X0,X2)))) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X0)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [rectify 129]
131. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X4,X3),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k2_relset_1(X0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X5,X3),X2) & m1_subset_1(X5,X1)) | r2_hidden(X3,k2_relset_1(X0,X2)))) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X0)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X4,X3),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k2_relset_1(sK0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X5,X3),X2) & m1_subset_1(X5,X1)) | r2_hidden(X3,k2_relset_1(sK0,X2)))) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,sK0)))) & ~v1_xboole_0(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
132. ? [X1] : (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X4,X3),X2) | ~m1_subset_1(X4,X1)) | ~r2_hidden(X3,k2_relset_1(sK0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X5,X3),X2) & m1_subset_1(X5,X1)) | r2_hidden(X3,k2_relset_1(sK0,X2)))) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,sK0)))) & ~v1_xboole_0(X1)) => (? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X4,X3),X2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(X3,k2_relset_1(sK0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X5,X3),X2) & m1_subset_1(X5,sK1)) | r2_hidden(X3,k2_relset_1(sK0,X2)))) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK1,sK0)))) & ~v1_xboole_0(sK1)) [choice axiom]
133. ? [X2] : (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X4,X3),X2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(X3,k2_relset_1(sK0,X2))) & (? [X5] : (r2_hidden(k4_tarski(X5,X3),X2) & m1_subset_1(X5,sK1)) | r2_hidden(X3,k2_relset_1(sK0,X2)))) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(sK1,sK0)))) => (? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X4,X3),sK2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(X3,k2_relset_1(sK0,sK2))) & (? [X5] : (r2_hidden(k4_tarski(X5,X3),sK2) & m1_subset_1(X5,sK1)) | r2_hidden(X3,k2_relset_1(sK0,sK2)))) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK1,sK0)))) [choice axiom]
134. ? [X3] : ((! [X4] : (~r2_hidden(k4_tarski(X4,X3),sK2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(X3,k2_relset_1(sK0,sK2))) & (? [X5] : (r2_hidden(k4_tarski(X5,X3),sK2) & m1_subset_1(X5,sK1)) | r2_hidden(X3,k2_relset_1(sK0,sK2)))) => ((! [X4] : (~r2_hidden(k4_tarski(X4,sK3),sK2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(sK3,k2_relset_1(sK0,sK2))) & (? [X5] : (r2_hidden(k4_tarski(X5,sK3),sK2) & m1_subset_1(X5,sK1)) | r2_hidden(sK3,k2_relset_1(sK0,sK2)))) [choice axiom]
135. ? [X5] : (r2_hidden(k4_tarski(X5,sK3),sK2) & m1_subset_1(X5,sK1)) => (r2_hidden(k4_tarski(sK4,sK3),sK2) & m1_subset_1(sK4,sK1)) [choice axiom]
136. ((((! [X4] : (~r2_hidden(k4_tarski(X4,sK3),sK2) | ~m1_subset_1(X4,sK1)) | ~r2_hidden(sK3,k2_relset_1(sK0,sK2))) & ((r2_hidden(k4_tarski(sK4,sK3),sK2) & m1_subset_1(sK4,sK1)) | r2_hidden(sK3,k2_relset_1(sK0,sK2)))) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK1,sK0)))) & ~v1_xboole_0(sK1)) & ~v1_xboole_0(sK0) [skolemisation 130,135,134,133,132,131]
141. ! [X0,X1] : ((k10_xtuple_0(X0) = X1 | ? [X2] : ((! [X3] : ~r2_hidden(k4_tarski(X3,X2),X0) | ~r2_hidden(X2,X1)) & (? [X3] : r2_hidden(k4_tarski(X3,X2),X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ! [X3] : ~r2_hidden(k4_tarski(X3,X2),X0)) & (? [X3] : r2_hidden(k4_tarski(X3,X2),X0) | ~r2_hidden(X2,X1))) | k10_xtuple_0(X0) != X1)) [nnf transformation 68]
142. ! [X0,X1] : ((k10_xtuple_0(X0) = X1 | ? [X2] : ((! [X3] : ~r2_hidden(k4_tarski(X3,X2),X0) | ~r2_hidden(X2,X1)) & (? [X4] : r2_hidden(k4_tarski(X4,X2),X0) | r2_hidden(X2,X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : ~r2_hidden(k4_tarski(X6,X5),X0)) & (? [X7] : r2_hidden(k4_tarski(X7,X5),X0) | ~r2_hidden(X5,X1))) | k10_xtuple_0(X0) != X1)) [rectify 141]
143. ! [X1,X0] : (? [X2] : ((! [X3] : ~r2_hidden(k4_tarski(X3,X2),X0) | ~r2_hidden(X2,X1)) & (? [X4] : r2_hidden(k4_tarski(X4,X2),X0) | r2_hidden(X2,X1))) => ((! [X3] : ~r2_hidden(k4_tarski(X3,sK7(X0,X1)),X0) | ~r2_hidden(sK7(X0,X1),X1)) & (? [X4] : r2_hidden(k4_tarski(X4,sK7(X0,X1)),X0) | r2_hidden(sK7(X0,X1),X1)))) [choice axiom]
144. ! [X1,X0] : (? [X4] : r2_hidden(k4_tarski(X4,sK7(X0,X1)),X0) => r2_hidden(k4_tarski(sK8(X0,X1),sK7(X0,X1)),X0)) [choice axiom]
145. ! [X5,X0] : (? [X7] : r2_hidden(k4_tarski(X7,X5),X0) => r2_hidden(k4_tarski(sK9(X0,X5),X5),X0)) [choice axiom]
146. ! [X0,X1] : ((k10_xtuple_0(X0) = X1 | ((! [X3] : ~r2_hidden(k4_tarski(X3,sK7(X0,X1)),X0) | ~r2_hidden(sK7(X0,X1),X1)) & (r2_hidden(k4_tarski(sK8(X0,X1),sK7(X0,X1)),X0) | r2_hidden(sK7(X0,X1),X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : ~r2_hidden(k4_tarski(X6,X5),X0)) & (r2_hidden(k4_tarski(sK9(X0,X5),X5),X0) | ~r2_hidden(X5,X1))) | k10_xtuple_0(X0) != X1)) [skolemisation 142,145,144,143]
147. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | (~r2_hidden(X1,X3) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [nnf transformation 74]
148. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [flattening 147]
167. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK1,sK0))) [cnf transformation 136] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
169. r2_hidden(k4_tarski(sK4,sK3),sK2) | r2_hidden(sK3,k2_relset_1(sK0,sK2)) [cnf transformation 136] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
170. ~r2_hidden(k4_tarski(X4,sK3),sK2) | ~m1_subset_1(X4,sK1) | ~r2_hidden(sK3,k2_relset_1(sK0,sK2)) [cnf transformation 136] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
182. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 89] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
185. v1_xboole_0(k10_xtuple_0(X0)) | ~v1_xboole_0(X0) [cnf transformation 92] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
186. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 93] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
198. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 105] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
199. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 107] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
200. r2_hidden(k4_tarski(sK9(X0,X5),X5),X0) | ~r2_hidden(X5,X1) | k10_xtuple_0(X0) != X1 [cnf transformation 146] {a:0,w:14,wCS:5600,sine:2,thAx:0,allAx:4,thDist:-4}
201. r2_hidden(X5,X1) | ~r2_hidden(k4_tarski(X6,X5),X0) | k10_xtuple_0(X0) != X1 [cnf transformation 146] {a:0,w:12,wCS:4800,sine:2,thAx:0,allAx:4,thDist:-4}
206. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 110] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
207. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1) [cnf transformation 112] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
209. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X0,X2) [cnf transformation 148] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
219. v1_relat_1(k2_zfmisc_1(X0,X1)) [cnf transformation 42] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
227. k2_relset_1(X0,X1) = k10_xtuple_0(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 122] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
232. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v5_relat_1(X2,X1) [cnf transformation 128] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
242. ~r2_hidden(k4_tarski(X6,X5),X0) | r2_hidden(X5,k10_xtuple_0(X0)) [equality resolution 201] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:4,thDist:-4}
243. r2_hidden(k4_tarski(sK9(X0,X5),X5),X0) | ~r2_hidden(X5,k10_xtuple_0(X0)) [equality resolution 200] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:4,thDist:-4}
245. 1 <=> r2_hidden(sK3,k2_relset_1(sK0,sK2)) [avatar definition]
246. r2_hidden(sK3,k2_relset_1(sK0,sK2)) <- (1) [avatar component clause 245] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
247. ~r2_hidden(sK3,k2_relset_1(sK0,sK2)) <- (~1) [avatar component clause 245] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
249. 2 <=> ! [X4] : (~r2_hidden(k4_tarski(X4,sK3),sK2) | ~m1_subset_1(X4,sK1)) [avatar definition]
250. ~r2_hidden(k4_tarski(X4,sK3),sK2) | ~m1_subset_1(X4,sK1) <- (2) [avatar component clause 249] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
251. ~1 | 2 [avatar split clause 170,249,245]
253. 3 <=> r2_hidden(k4_tarski(sK4,sK3),sK2) [avatar definition]
255. r2_hidden(k4_tarski(sK4,sK3),sK2) <- (3) [avatar component clause 253] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
256. 1 | 3 [avatar split clause 169,253,245]
263. 5 <=> m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK1,sK0))) [avatar definition]
265. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK1,sK0))) <- (5) [avatar component clause 263] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
266. 5 [avatar split clause 167,263]
337. 18 <=> v1_xboole_0(sK2) [avatar definition]
355. 21 <=> v1_xboole_0(k2_relset_1(sK0,sK2)) [avatar definition]
357. ~v1_xboole_0(k2_relset_1(sK0,sK2)) <- (~21) [avatar component clause 355] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
368. ~v1_xboole_0(k2_relset_1(sK0,sK2)) <- (1) [resolution 246,206] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
370. ~21 | ~1 [avatar split clause 368,245,355]
381. v1_relat_1(sK2) | ~v1_relat_1(k2_zfmisc_1(sK1,sK0)) <- (5) [resolution 182,265] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
387. v1_relat_1(sK2) <- (5) [subsumption resolution 381,219] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
389. 23 <=> v1_relat_1(sK2) [avatar definition]
391. v1_relat_1(sK2) <- (23) [avatar component clause 389] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
392. 23 | ~5 [avatar split clause 387,263,389]
394. v1_xboole_0(sK2) | ~v1_xboole_0(k2_zfmisc_1(sK1,sK0)) <- (5) [resolution 186,265] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
402. 24 <=> v1_xboole_0(k2_zfmisc_1(sK1,sK0)) [avatar definition]
436. ~24 | 18 | ~5 [avatar split clause 394,263,337,402]
495. v5_relat_1(sK2,sK0) <- (5) [resolution 232,265] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
502. 30 <=> v5_relat_1(sK2,sK0) [avatar definition]
504. v5_relat_1(sK2,sK0) <- (30) [avatar component clause 502] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
505. 30 | ~5 [avatar split clause 495,263,502]
516. m1_subset_1(X3,k2_zfmisc_1(sK1,sK0)) | ~r2_hidden(X3,sK2) <- (5) [resolution 207,265] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
522. ~r2_hidden(X0,sK2) | v1_xboole_0(k2_zfmisc_1(sK1,sK0)) | r2_hidden(X0,k2_zfmisc_1(sK1,sK0)) <- (5) [resolution 516,199] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
573. 33 <=> ! [X0] : (~r2_hidden(X0,sK2) | r2_hidden(X0,k2_zfmisc_1(sK1,sK0))) [avatar definition]
574. r2_hidden(X0,k2_zfmisc_1(sK1,sK0)) | ~r2_hidden(X0,sK2) <- (33) [avatar component clause 573] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
575. 24 | 33 | ~5 [avatar split clause 522,263,573,402]
583. ~v1_xboole_0(k10_xtuple_0(sK2)) | ~v5_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~21) [superposition 357,227] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
599. ~v1_xboole_0(k10_xtuple_0(sK2)) | ~v1_relat_1(sK2) <- (~21, 30) [subsumption resolution 583,504] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
600. ~v1_xboole_0(k10_xtuple_0(sK2)) <- (~21, 23, 30) [subsumption resolution 599,391] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
602. 37 <=> v1_xboole_0(k10_xtuple_0(sK2)) [avatar definition]
604. ~v1_xboole_0(k10_xtuple_0(sK2)) <- (~37) [avatar component clause 602] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
605. ~37 | 21 | ~23 | ~30 [avatar split clause 600,502,389,355,602]
609. 38 <=> r2_hidden(sK3,k10_xtuple_0(sK2)) [avatar definition]
611. r2_hidden(sK3,k10_xtuple_0(sK2)) <- (38) [avatar component clause 609] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
613. ~r2_hidden(sK3,k10_xtuple_0(sK2)) | ~v5_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (~1) [superposition 247,227] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
614. ~r2_hidden(sK3,k10_xtuple_0(sK2)) | ~v1_relat_1(sK2) <- (~1, 30) [subsumption resolution 613,504] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
615. ~r2_hidden(sK3,k10_xtuple_0(sK2)) <- (~1, 23, 30) [subsumption resolution 614,391] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
616. ~38 | 1 | ~23 | ~30 [avatar split clause 615,502,389,245,609]
617. r2_hidden(sK3,k10_xtuple_0(sK2)) <- (3) [resolution 255,242] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
621. 38 | ~3 [avatar split clause 617,253,609]
631. r2_hidden(sK3,k10_xtuple_0(sK2)) | ~v5_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (1) [superposition 246,227] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
633. r2_hidden(sK3,k10_xtuple_0(sK2)) | ~v1_relat_1(sK2) <- (1, 30) [subsumption resolution 631,504] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
634. r2_hidden(sK3,k10_xtuple_0(sK2)) <- (1, 23, 30) [subsumption resolution 633,391] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
635. 38 | ~1 | ~23 | ~30 [avatar split clause 634,502,389,245,609]
637. ~v1_xboole_0(sK2) <- (~37) [resolution 604,185] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
641. ~18 | 37 [avatar split clause 637,602,337]
660. ~r2_hidden(sK3,k10_xtuple_0(sK2)) | ~m1_subset_1(sK9(sK2,sK3),sK1) <- (2) [resolution 243,250] {a:1,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
665. 40 <=> m1_subset_1(sK9(sK2,sK3),sK1) [avatar definition]
667. ~m1_subset_1(sK9(sK2,sK3),sK1) <- (~40) [avatar component clause 665] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
668. ~40 | ~38 | ~2 [avatar split clause 660,249,609,665]
715. ~r2_hidden(k4_tarski(X5,X6),sK2) | r2_hidden(X5,sK1) <- (33) [resolution 574,209] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
786. r2_hidden(sK9(sK2,X0),sK1) | ~r2_hidden(X0,k10_xtuple_0(sK2)) <- (33) [resolution 715,243] {a:4,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
791. m1_subset_1(sK9(sK2,X1),sK1) | ~r2_hidden(X1,k10_xtuple_0(sK2)) <- (33) [resolution 786,198] {a:5,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
800. ~r2_hidden(sK3,k10_xtuple_0(sK2)) <- (33, ~40) [resolution 791,667] {a:6,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:24,thDist:-24}
803. $false <- (33, 38, ~40) [subsumption resolution 800,611] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:24,thDist:-24}
804. ~33 | ~38 | 40 [avatar contradiction clause 803]
805. $false [avatar sat refutation 251,256,266,370,392,436,505,575,605,616,621,635,641,668,804] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:707,thDist:-707}
% SZS output end Proof for relset_1__t25_relset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 63
% >>> Preprocessing
% Introduced skolems: 18
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 77
% Generated clauses: 468
% Active clauses: 152
% Passive clauses: 269
% Final active clauses: 132
% Final passive clauses: 108
% Main loop iterations started: 154
% 
% >>> Simplifying Inferences
% Duplicate literals: 7
% Fw subsumption resolutions: 40
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 21
% Forward subsumptions: 68
% 
% >>> Generating Inferences
% Binary resolution: 241
% Unit resulting resolution: 3
% Forward superposition: 10
% Backward superposition: 11
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 10
% Split components: 20
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 68
% SAT solver unit clauses: 11
% SAT solver binary clauses: 31
% 
% Memory used [KB]: 5500
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% splitting model update: 0.001 s ( own 0.0000 s ) 
% backward subsumption index maintenance: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

