i: [1,1017,0,14,3,1,0,0] t4_glib_003
a: [3,1470,0,2,1,1017]
s: 1470
i: [1,1016,0,14,3,1,0,0] t4_glib_003
a: [3,1475,0,2,1,1016]
s: 1475
i: [1,1015,0,14,3,1,0,0] t4_glib_003
a: [3,1480,0,2,1,1015]
s: 1480
i: [1,1014,0,14,3,1,0,0] t4_glib_003
a: [3,1485,0,2,1,1014]
s: 1485
i: [1,1012,0,14,3,1,0,0] t4_glib_003
a: [3,1495,0,2,1,1012]
s: 1495
i: [1,1013,0,15,3,1,0,0] t4_glib_003
a: [3,1490,0,3,1,1013]
s: 1490
i: [1,1018,0,18,3,1,0,0] t4_glib_003
a: [3,1457,0,6,1,1018]
s: 1457
i: [1,1019,0,2,1,0,0,255] spc7_boole
a: [3,1884,0,2,1,1019]
s: 1884
i: [1,1020,0,2,1,0,0,255] spc6_boole
a: [3,1889,0,2,1,1020]
s: 1889
i: [1,1021,0,2,1,0,0,255] spc5_boole
a: [3,1894,0,2,1,1021]
s: 1894
i: [1,1022,0,2,1,0,0,255] spc2_boole
a: [3,1899,0,2,1,1022]
s: 1899
i: [1,1023,0,2,1,0,0,255] spc1_boole
a: [3,1904,0,2,1,1023]
s: 1904
i: [1,1024,0,2,1,0,0,255] fc3_numbers
a: [3,1909,0,2,1,1024]
s: 1909
i: [1,1025,0,2,1,0,0,255] fc4_numbers
a: [3,1914,0,2,1,1025]
s: 1914
i: [1,1026,0,2,1,0,0,2] fc1_numbers
a: [3,1919,0,2,1,1026]
s: 1919
i: [1,1027,0,2,1,0,0,255] fc7_card_1
a: [3,1924,0,2,1,1027]
s: 1924
i: [1,1028,0,2,1,0,0,255] fc7_numbers
a: [3,1929,0,2,1,1028]
s: 1929
i: [1,1029,0,2,1,0,0,255] fc6_numbers
a: [3,1934,0,2,1,1029]
s: 1934
i: [1,1030,0,2,1,0,0,2] fc8_numbers
a: [3,1939,0,2,1,1030]
s: 1939
i: [1,1031,0,2,1,0,0,255] fc4_membered
a: [3,1944,0,2,1,1031]
s: 1944
i: [1,1032,0,2,1,0,0,255] fc5_membered
a: [3,1949,0,2,1,1032]
s: 1949
i: [1,1033,0,2,1,0,0,255] fc3_membered
a: [3,1954,0,2,1,1033]
s: 1954
i: [1,1034,0,2,1,0,0,255] fc4_card_1
a: [3,1959,0,2,1,1034]
s: 1959
i: [1,1035,0,2,1,0,0,255] fc6_membered
a: [3,1964,0,2,1,1035]
s: 1964
i: [1,1036,0,2,1,0,0,255] fc1_xboole_0
a: [3,1969,0,2,1,1036]
s: 1969
i: [1,1052,0,2,1,0,0,255] fc6_ordinal1
a: [3,2032,0,2,1,1052]
s: 2032
i: [1,1051,0,2,1,0,0,255] fc6_ordinal1
a: [3,2037,0,2,1,1051]
s: 2037
i: [1,1327,0,2,1,0,0,2] rc2_xboole_0
a: [3,2117,0,2,1,1327]
s: 2117
i: [1,1328,0,2,1,0,0,1] rc3_card_1
a: [3,2122,0,2,1,1328]
s: 2122
i: [1,1329,0,2,1,0,0,255] rc1_card_1
a: [3,2127,0,2,1,1329]
s: 2127
i: [1,1330,0,2,1,0,0,255] rc4_ordinal1
a: [3,2132,0,2,1,1330]
s: 2132
i: [1,1331,0,2,1,0,0,255] rc1_ordinal1
a: [3,2137,0,2,1,1331]
s: 2137
i: [1,1332,0,2,1,0,0,2] rc1_xboole_0
a: [3,2142,0,2,1,1332]
s: 2142
i: [1,1334,0,2,1,0,0,255] rc8_finseq_1
a: [3,2147,0,2,1,1334]
s: 2147
i: [1,1333,0,2,1,0,0,255] rc8_finseq_1
a: [3,2152,0,2,1,1333]
s: 2152
i: [1,1336,0,2,1,0,0,255] rc7_funct_1
a: [3,2157,0,2,1,1336]
s: 2157
i: [1,1335,0,2,1,0,0,255] rc7_funct_1
a: [3,2162,0,2,1,1335]
s: 2162
i: [1,1340,0,2,1,0,0,255] rc2_ordinal1
a: [3,2167,0,2,1,1340]
s: 2167
i: [1,1339,0,2,1,0,0,255] rc2_ordinal1
a: [3,2172,0,2,1,1339]
s: 2172
i: [1,1338,0,2,1,0,0,255] rc2_ordinal1
a: [3,2177,0,2,1,1338]
s: 2177
i: [1,1337,0,2,1,0,0,255] rc2_ordinal1
a: [3,2182,0,2,1,1337]
s: 2182
i: [1,1347,0,2,1,0,0,255] rc1_nat_1
a: [3,2187,0,2,1,1347]
s: 2187
i: [1,1346,0,2,1,0,0,255] rc1_nat_1
a: [3,2192,0,2,1,1346]
s: 2192
i: [1,1345,0,2,1,0,0,255] rc1_nat_1
a: [3,2197,0,2,1,1345]
s: 2197
i: [1,1344,0,2,1,0,0,255] rc1_nat_1
a: [3,2202,0,2,1,1344]
s: 2202
i: [1,1343,0,2,1,0,0,255] rc1_nat_1
a: [3,2207,0,2,1,1343]
s: 2207
i: [1,1342,0,2,1,0,0,255] rc1_nat_1
a: [3,2212,0,2,1,1342]
s: 2212
i: [1,1341,0,2,1,0,0,255] rc1_nat_1
a: [3,2217,0,2,1,1341]
s: 2217
i: [1,1349,0,2,1,0,0,255] rc1_membered
a: [3,2222,0,2,1,1349]
s: 2222
i: [1,1348,0,2,1,0,0,255] rc1_membered
a: [3,2227,0,2,1,1348]
s: 2227
i: [1,1351,0,2,1,0,0,255] rc2_membered
a: [3,2232,0,2,1,1351]
s: 2232
i: [1,1350,0,2,1,0,0,255] rc2_membered
a: [3,2237,0,2,1,1350]
s: 2237
i: [1,1353,0,2,1,0,0,255] rc3_membered
a: [3,2242,0,2,1,1353]
s: 2242
i: [1,1352,0,2,1,0,0,255] rc3_membered
a: [3,2247,0,2,1,1352]
s: 2247
i: [1,1355,0,2,1,0,0,255] rc5_ordinal1
a: [3,2252,0,2,1,1355]
s: 2252
i: [1,1354,0,2,1,0,0,255] rc5_ordinal1
a: [3,2257,0,2,1,1354]
s: 2257
i: [1,1357,0,2,1,0,0,1] rc1_finset_1
a: [3,2262,0,2,1,1357]
s: 2262
i: [1,1356,0,2,1,0,0,1] rc1_finset_1
a: [3,2267,0,2,1,1356]
s: 2267
i: [1,1360,0,2,1,0,0,255] rc8_finset_1
a: [3,2272,0,2,1,1360]
s: 2272
i: [1,1359,0,2,1,0,0,255] rc8_finset_1
a: [3,2277,0,2,1,1359]
s: 2277
i: [1,1358,0,2,1,0,0,255] rc8_finset_1
a: [3,2282,0,2,1,1358]
s: 2282
i: [1,1362,0,2,1,0,0,2] rc1_relat_1
a: [3,2287,0,2,1,1362]
s: 2287
i: [1,1361,0,2,1,0,0,2] rc1_relat_1
a: [3,2292,0,2,1,1361]
s: 2292
i: [1,1365,0,2,1,0,0,255] rc4_funct_1
a: [3,2297,0,2,1,1365]
s: 2297
i: [1,1364,0,2,1,0,0,255] rc4_funct_1
a: [3,2302,0,2,1,1364]
s: 2302
i: [1,1363,0,2,1,0,0,255] rc4_funct_1
a: [3,2307,0,2,1,1363]
s: 2307
i: [1,1369,0,2,1,0,0,255] rc9_finset_1
a: [3,2312,0,2,1,1369]
s: 2312
i: [1,1368,0,2,1,0,0,255] rc9_finset_1
a: [3,2317,0,2,1,1368]
s: 2317
i: [1,1367,0,2,1,0,0,255] rc9_finset_1
a: [3,2322,0,2,1,1367]
s: 2322
i: [1,1366,0,2,1,0,0,255] rc9_finset_1
a: [3,2327,0,2,1,1366]
s: 2327
i: [1,1373,0,2,1,0,0,1] rc3_finset_1
a: [3,2332,0,2,1,1373]
s: 2332
i: [1,1372,0,2,1,0,0,1] rc3_finset_1
a: [3,2337,0,2,1,1372]
s: 2337
i: [1,1371,0,2,1,0,0,1] rc3_finset_1
a: [3,2342,0,2,1,1371]
s: 2342
i: [1,1370,0,2,1,0,0,1] rc3_finset_1
a: [3,2347,0,2,1,1370]
s: 2347
i: [1,1380,0,2,1,0,0,255] rc2_valued_0
a: [3,2352,0,2,1,1380]
s: 2352
i: [1,1379,0,2,1,0,0,255] rc2_valued_0
a: [3,2357,0,2,1,1379]
s: 2357
i: [1,1378,0,2,1,0,0,255] rc2_valued_0
a: [3,2362,0,2,1,1378]
s: 2362
i: [1,1375,0,2,1,0,0,255] rc2_valued_0
a: [3,2377,0,2,1,1375]
s: 2377
i: [1,1374,0,2,1,0,0,255] rc2_valued_0
a: [3,2382,0,2,1,1374]
s: 2382
i: [1,1385,0,2,1,0,0,255] rc2_card_1
a: [3,2387,0,2,1,1385]
s: 2387
i: [1,1384,0,2,1,0,0,255] rc2_card_1
a: [3,2392,0,2,1,1384]
s: 2392
i: [1,1383,0,2,1,0,0,255] rc2_card_1
a: [3,2397,0,2,1,1383]
s: 2397
i: [1,1382,0,2,1,0,0,255] rc2_card_1
a: [3,2402,0,2,1,1382]
s: 2402
i: [1,1381,0,2,1,0,0,255] rc2_card_1
a: [3,2407,0,2,1,1381]
s: 2407
i: [1,1386,0,2,1,0,0,255] rc3_xreal_0
a: [3,2412,0,2,1,1386]
s: 2412
i: [1,1387,0,2,1,0,0,255] rc4_xreal_0
a: [3,2417,0,2,1,1387]
s: 2417
i: [1,1388,0,2,1,0,0,255] rc2_relat_1
a: [3,2422,0,2,1,1388]
s: 2422
i: [1,1390,0,2,1,0,0,255] rc3_funct_1
a: [3,2427,0,2,1,1390]
s: 2427
i: [1,1389,0,2,1,0,0,255] rc3_funct_1
a: [3,2432,0,2,1,1389]
s: 2432
i: [1,1397,0,2,1,0,0,255] rc5_finseq_1
a: [3,2437,0,2,1,1397]
s: 2437
i: [1,1396,0,2,1,0,0,255] rc5_finseq_1
a: [3,2442,0,2,1,1396]
s: 2442
i: [1,1395,0,2,1,0,0,255] rc5_finseq_1
a: [3,2447,0,2,1,1395]
s: 2447
i: [1,1394,0,2,1,0,0,255] rc5_finseq_1
a: [3,2452,0,2,1,1394]
s: 2452
i: [1,1393,0,2,1,0,0,255] rc5_finseq_1
a: [3,2457,0,2,1,1393]
s: 2457
i: [1,1391,0,2,1,0,0,255] rc5_finseq_1
a: [3,2467,0,2,1,1391]
s: 2467
i: [1,1399,0,2,1,0,0,1] rc1_funct_1
a: [3,2472,0,2,1,1399]
s: 2472
i: [1,1398,0,2,1,0,0,1] rc1_funct_1
a: [3,2477,0,2,1,1398]
s: 2477
i: [1,1402,0,2,1,0,0,255] rc5_funct_1
a: [3,2482,0,2,1,1402]
s: 2482
i: [1,1401,0,2,1,0,0,255] rc5_funct_1
a: [3,2487,0,2,1,1401]
s: 2487
i: [1,1400,0,2,1,0,0,255] rc5_funct_1
a: [3,2492,0,2,1,1400]
s: 2492
i: [1,1404,0,2,1,0,0,255] rc2_funct_1
a: [3,2497,0,2,1,1404]
s: 2497
i: [1,1403,0,2,1,0,0,255] rc2_funct_1
a: [3,2502,0,2,1,1403]
s: 2502
i: [1,1407,0,2,1,0,0,255] rc1_funcop_1
a: [3,2507,0,2,1,1407]
s: 2507
i: [1,1406,0,2,1,0,0,255] rc1_funcop_1
a: [3,2512,0,2,1,1406]
s: 2512
i: [1,1405,0,2,1,0,0,255] rc1_funcop_1
a: [3,2517,0,2,1,1405]
s: 2517
i: [1,1411,0,2,1,0,0,255] rc2_funcop_1
a: [3,2522,0,2,1,1411]
s: 2522
i: [1,1410,0,2,1,0,0,255] rc2_funcop_1
a: [3,2527,0,2,1,1410]
s: 2527
i: [1,1409,0,2,1,0,0,255] rc2_funcop_1
a: [3,2532,0,2,1,1409]
s: 2532
i: [1,1408,0,2,1,0,0,255] rc2_funcop_1
a: [3,2537,0,2,1,1408]
s: 2537
i: [1,1414,0,2,1,0,0,255] rc1_valued_0
a: [3,2542,0,2,1,1414]
s: 2542
i: [1,1413,0,2,1,0,0,255] rc1_valued_0
a: [3,2547,0,2,1,1413]
s: 2547
i: [1,1412,0,2,1,0,0,255] rc1_valued_0
a: [3,2552,0,2,1,1412]
s: 2552
i: [1,1417,0,2,1,0,0,255] rc3_finseq_1
a: [3,2557,0,2,1,1417]
s: 2557
i: [1,1416,0,2,1,0,0,255] rc3_finseq_1
a: [3,2562,0,2,1,1416]
s: 2562
i: [1,1415,0,2,1,0,0,255] rc3_finseq_1
a: [3,2567,0,2,1,1415]
s: 2567
i: [1,1424,0,2,1,0,0,255] rc4_finseq_1
a: [3,2572,0,2,1,1424]
s: 2572
i: [1,1423,0,2,1,0,0,255] rc4_finseq_1
a: [3,2577,0,2,1,1423]
s: 2577
i: [1,1422,0,2,1,0,0,255] rc4_finseq_1
a: [3,2582,0,2,1,1422]
s: 2582
i: [1,1421,0,2,1,0,0,255] rc4_finseq_1
a: [3,2587,0,2,1,1421]
s: 2587
i: [1,1420,0,2,1,0,0,255] rc4_finseq_1
a: [3,2592,0,2,1,1420]
s: 2592
i: [1,1418,0,2,1,0,0,255] rc4_finseq_1
a: [3,2602,0,2,1,1418]
s: 2602
i: [1,1428,0,2,1,0,0,1] rc1_glib_000
a: [3,2607,0,2,1,1428]
s: 2607
i: [1,1427,0,2,1,0,0,1] rc1_glib_000
a: [3,2612,0,2,1,1427]
s: 2612
i: [1,1425,0,2,1,0,0,1] rc1_glib_000
a: [3,2622,0,2,1,1425]
s: 2622
i: [1,1433,0,2,1,0,0,1] rc2_glib_000
a: [3,2627,0,2,1,1433]
s: 2627
i: [1,1432,0,2,1,0,0,1] rc2_glib_000
a: [3,2632,0,2,1,1432]
s: 2632
i: [1,1431,0,2,1,0,0,1] rc2_glib_000
a: [3,2637,0,2,1,1431]
s: 2637
i: [1,1429,0,2,1,0,0,1] rc2_glib_000
a: [3,2647,0,2,1,1429]
s: 2647
i: [1,1441,0,2,1,0,0,1] rc1_glib_003
a: [3,2652,0,2,1,1441]
s: 2652
i: [1,1440,0,2,1,0,0,1] rc1_glib_003
a: [3,2657,0,2,1,1440]
s: 2657
i: [1,1439,0,2,1,0,0,1] rc1_glib_003
a: [3,2662,0,2,1,1439]
s: 2662
i: [1,1438,0,2,1,0,0,1] rc1_glib_003
a: [3,2667,0,2,1,1438]
s: 2667
i: [1,1437,0,2,1,0,0,1] rc1_glib_003
a: [3,2672,0,2,1,1437]
s: 2672
i: [1,1436,0,2,1,0,0,1] rc1_glib_003
a: [3,2677,0,2,1,1436]
s: 2677
i: [1,1434,0,2,1,0,0,1] rc1_glib_003
a: [3,2687,0,2,1,1434]
s: 2687
i: [1,1450,0,2,1,0,0,255] rc3_nat_1
a: [3,2697,0,2,1,1450]
s: 2697
i: [1,1449,0,2,1,0,0,255] rc3_nat_1
a: [3,2702,0,2,1,1449]
s: 2702
i: [1,1448,0,2,1,0,0,255] rc3_nat_1
a: [3,2707,0,2,1,1448]
s: 2707
i: [1,1447,0,2,1,0,0,255] rc3_nat_1
a: [3,2712,0,2,1,1447]
s: 2712
i: [1,1446,0,2,1,0,0,255] rc3_nat_1
a: [3,2717,0,2,1,1446]
s: 2717
i: [1,1445,0,2,1,0,0,255] rc3_nat_1
a: [3,2722,0,2,1,1445]
s: 2722
i: [1,1444,0,2,1,0,0,255] rc3_nat_1
a: [3,2727,0,2,1,1444]
s: 2727
i: [1,1453,0,2,1,0,0,255] rc2_nat_1
a: [3,2737,0,2,1,1453]
s: 2737
i: [1,1452,0,2,1,0,0,255] rc2_nat_1
a: [3,2742,0,2,1,1452]
s: 2742
i: [1,1273,0,3,1,0,0,255] rc9_funct_1
s: 1273
i: [1,1275,0,3,1,0,0,255] rc9_funct_1
s: 1275
i: [1,1276,0,3,1,0,0,255] rc4_funcop_1
s: 1276
i: [1,1278,0,3,1,0,0,255] rc4_funcop_1
s: 1278
i: [1,1279,0,3,1,0,0,255] rc9_finseq_1
s: 1279
i: [1,1282,0,3,1,0,0,255] rc9_finseq_1
s: 1282
i: [1,1283,0,3,1,0,0,255] rc9_finseq_1
s: 1283
i: [1,1284,0,3,1,0,0,255] rc9_finseq_1
s: 1284
i: [1,1285,0,3,1,0,0,255] rc9_finseq_1
s: 1285
i: [1,1286,0,3,1,0,0,255] rc3_ordinal1
s: 1286
i: [1,1288,0,3,1,0,0,255] rc3_ordinal1
s: 1288
i: [1,1289,0,3,1,0,0,255] rc6_valued_0
s: 1289
i: [1,1291,0,3,1,0,0,255] rc6_valued_0
s: 1291
i: [1,1292,0,3,1,0,0,255] rc6_valued_0
s: 1292
i: [1,1293,0,3,1,0,0,255] reflexivity_r1_tarski
s: 1293
i: [1,1037,0,3,1,0,0,255] d1_glib_003
a: [3,1974,0,3,1,1037]
s: 1974
r: 1894
d: [2,2749,0,2,1,2,23,1894,1974]
a: [3,2753,0,2,1,2749]
s: 2753
i: [1,1038,0,3,1,0,0,255] d2_glib_003
a: [3,1979,0,3,1,1038]
s: 1979
r: 1889
d: [2,2755,0,2,1,2,23,1889,1979]
a: [3,2759,0,2,1,2755]
s: 2759
i: [1,1039,0,3,1,0,0,255] d3_glib_003
a: [3,1984,0,3,1,1039]
s: 1984
r: 1884
d: [2,2761,0,2,1,2,23,1884,1984]
a: [3,2765,0,2,1,2761]
s: 2765
i: [1,1040,0,3,1,0,0,255] d2_glib_000
a: [3,1989,0,3,1,1040]
s: 1989
i: [1,1041,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,1994,0,3,1,1041]
s: 1994
r: 2037
r: 2032
r: 1964
r: 1959
r: 1924
d: [2,2767,0,2,1,2,23,2037,1994]
a: [3,2775,0,2,1,2767]
s: 2775
d: [2,2768,0,2,1,2,23,2032,1994]
a: [3,2780,0,2,1,2768]
s: 2780
d: [2,2769,0,2,1,2,23,1964,1994]
a: [3,2785,0,2,1,2769]
s: 2785
d: [2,2770,0,2,1,2,23,1959,1994]
a: [3,2790,0,2,1,2770]
s: 2790
d: [2,2771,0,2,1,2,23,1924,1994]
a: [3,2795,0,2,1,2771]
s: 2795
i: [1,1042,0,3,1,0,0,255] dt_k2_glib_000
a: [3,1999,0,3,1,1042]
s: 1999
i: [1,1043,0,3,1,0,0,255] t1_numerals
a: [3,2004,0,3,1,1043]
s: 2004
i: [1,1055,0,3,1,0,0,255] spc7_numerals
a: [3,2042,0,3,1,1055]
s: 2042
i: [1,1054,0,3,1,0,0,255] spc7_numerals
a: [3,2047,0,3,1,1054]
s: 2047
i: [1,1058,0,3,1,0,0,255] spc6_numerals
a: [3,2057,0,3,1,1058]
s: 2057
i: [1,1057,0,3,1,0,0,255] spc6_numerals
a: [3,2062,0,3,1,1057]
s: 2062
i: [1,1061,0,3,1,0,0,255] spc5_numerals
a: [3,2072,0,3,1,1061]
s: 2072
i: [1,1060,0,3,1,0,0,255] spc5_numerals
a: [3,2077,0,3,1,1060]
s: 2077
i: [1,1064,0,3,1,0,0,255] spc2_numerals
a: [3,2087,0,3,1,1064]
s: 2087
i: [1,1063,0,3,1,0,0,255] spc2_numerals
a: [3,2092,0,3,1,1063]
s: 2092
i: [1,1067,0,3,1,0,0,255] spc1_numerals
a: [3,2102,0,3,1,1067]
s: 2102
i: [1,1066,0,3,1,0,0,255] spc1_numerals
a: [3,2107,0,3,1,1066]
s: 2107
i: [1,1377,0,3,1,0,0,255] rc2_valued_0
a: [3,2367,0,3,1,1377]
s: 2367
i: [1,1376,0,3,1,0,0,255] rc2_valued_0
a: [3,2372,0,3,1,1376]
s: 2372
i: [1,1392,0,3,1,0,0,255] rc5_finseq_1
a: [3,2462,0,3,1,1392]
s: 2462
i: [1,1419,0,3,1,0,0,255] rc4_finseq_1
a: [3,2597,0,3,1,1419]
s: 2597
i: [1,1426,0,3,1,0,0,1] rc1_glib_000
a: [3,2617,0,3,1,1426]
s: 2617
i: [1,1430,0,3,1,0,0,1] rc2_glib_000
a: [3,2642,0,3,1,1430]
s: 2642
i: [1,1435,0,3,1,0,0,1] rc1_glib_003
a: [3,2682,0,3,1,1435]
s: 2682
i: [1,1442,0,3,1,0,0,255] rc1_int_1
a: [3,2692,0,3,1,1442]
s: 2692
i: [1,1443,0,3,1,0,0,255] rc3_nat_1
a: [3,2732,0,3,1,1443]
s: 2732
d: [2,2799,0,3,1,2,22,2004,1994]
a: [3,2803,0,3,1,2799]
s: 2803
d: [2,2805,0,3,1,2,22,2042,1984]
a: [3,2809,0,3,1,2805]
s: 2809
d: [2,2811,0,3,1,2,22,2047,1984]
a: [3,2815,0,3,1,2811]
s: 2815
d: [2,2817,0,3,1,2,22,2057,1979]
a: [3,2821,0,3,1,2817]
s: 2821
d: [2,2823,0,3,1,2,22,2062,1979]
a: [3,2827,0,3,1,2823]
s: 2827
d: [2,2829,0,3,1,2,22,2072,1974]
a: [3,2833,0,3,1,2829]
s: 2833
d: [2,2835,0,3,1,2,22,2077,1974]
a: [3,2839,0,3,1,2835]
s: 2839
i: [1,1076,0,4,2,0,0,255] cc5_finset_1
s: 1076
i: [1,1083,0,4,2,0,0,255] cc4_finset_1
s: 1083
i: [1,1084,0,4,2,0,0,255] cc9_finseq_1
s: 1084
d: [2,2841,1,2,1,1,40,1084,2147]
a: [3,2845,1,2,1,2841]
s: 2845
i: [1,1089,0,4,2,0,0,255] cc2_membered
s: 1089
i: [1,1090,0,4,2,0,0,255] cc3_membered
s: 1090
d: [2,2847,1,2,1,1,40,1089,1949]
a: [3,2851,1,2,1,2847]
s: 2851
i: [1,1093,0,4,2,0,0,255] cc5_membered
s: 1093
d: [2,2853,1,2,1,1,40,1090,1944]
a: [3,2857,1,2,1,2853]
s: 2857
i: [1,1094,0,4,2,0,0,255] cc4_membered
s: 1094
i: [1,1102,0,4,2,0,0,255] cc1_card_1
s: 1102
d: [2,2865,1,2,1,1,40,1093,1954]
a: [3,2869,1,2,1,2865]
s: 2869
i: [1,1107,0,4,2,0,0,255] cc1_membered
s: 1107
d: [2,2877,1,2,1,1,40,1094,1954]
a: [3,2882,1,2,1,2877]
s: 2882
i: [1,1110,0,4,2,0,0,255] cc3_card_1
s: 1110
i: [1,1111,0,4,2,0,0,255] cc6_ordinal1
s: 1111
d: [2,2890,1,2,1,1,40,1102,2127]
a: [3,2897,1,2,1,2890]
s: 2897
i: [1,1112,0,4,2,0,0,255] cc5_card_1
s: 1112
d: [2,2899,1,2,1,1,40,1107,2785]
a: [3,2906,1,2,1,2899]
s: 2906
i: [1,1114,0,4,2,0,0,255] cc3_nat_1
s: 1114
i: [1,1115,0,4,2,0,0,255] cc1_nat_1
s: 1115
d: [2,2900,1,2,1,1,40,1107,2222]
a: [3,2911,1,2,1,2900]
s: 2911
i: [1,1124,0,4,2,0,0,255] cc1_ordinal1
s: 1124
d: [2,2901,1,2,1,1,40,1107,2232]
a: [3,2916,1,2,1,2901]
s: 2916
i: [1,1125,0,4,2,0,0,255] cc1_ordinal1
s: 1125
i: [1,1127,0,4,2,0,0,255] cc9_ordinal1
s: 1127
d: [2,2902,1,2,1,1,40,1107,2242]
a: [3,2921,1,2,1,2902]
s: 2921
i: [1,1128,0,4,2,0,0,255] cc8_finseq_1
s: 1128
d: [2,2923,1,2,1,1,40,1110,2132]
a: [3,2930,1,2,1,2923]
s: 2930
i: [1,1129,0,4,2,0,0,255] cc6_finset_1
s: 1129
i: [1,1130,0,4,2,0,0,255] cc7_funct_1
s: 1130
d: [2,2925,1,2,1,1,40,1110,2252]
a: [3,2935,1,2,1,2925]
s: 2935
i: [1,1131,0,4,2,0,0,255] cc2_card_1
s: 1131
d: [2,2937,1,2,1,1,40,1111,2132]
a: [3,2944,1,2,1,2937]
s: 2944
i: [1,1132,0,4,2,0,0,255] cc18_membered
s: 1132
i: [1,1133,0,4,2,0,0,255] cc7_ordinal1
s: 1133
d: [2,2939,1,2,1,1,40,1111,2252]
a: [3,2949,1,2,1,2939]
s: 2949
i: [1,1134,0,4,2,0,0,255] cc3_ordinal1
s: 1134
d: [2,2951,1,2,1,1,40,1112,2132]
a: [3,2958,1,2,1,2951]
s: 2958
i: [1,1135,0,4,2,0,0,1] cc1_finset_1
s: 1135
i: [1,1136,0,4,2,0,0,1] cc1_funct_1
s: 1136
d: [2,2953,1,2,1,1,40,1112,2252]
a: [3,2963,1,2,1,2953]
s: 2963
i: [1,1137,0,4,2,0,0,2] cc1_relat_1
s: 1137
d: [2,2971,1,2,1,1,40,1114,2412]
a: [3,2975,1,2,1,2971]
s: 2975
i: [1,1272,0,4,1,0,0,2] existence_m1_subset_1
s: 1272
i: [1,1274,0,4,1,0,0,255] rc9_funct_1
s: 1274
d: [2,2983,1,2,1,1,40,1124,2780]
a: [3,2994,1,2,1,2983]
s: 2994
i: [1,1277,0,4,1,0,0,255] rc4_funcop_1
s: 1277
d: [2,2985,1,2,1,1,40,1124,2137]
a: [3,3004,1,2,1,2985]
s: 3004
i: [1,1280,0,4,1,0,0,255] rc9_finseq_1
s: 1280
i: [1,1281,0,4,1,0,0,255] rc9_finseq_1
s: 1281
d: [2,2990,1,2,1,1,40,1124,2737]
a: [3,3009,1,2,1,2990]
s: 3009
i: [1,1287,0,4,1,0,0,255] rc3_ordinal1
s: 1287
d: [2,3017,1,2,1,1,40,1125,2780]
a: [3,3028,1,2,1,3017]
s: 3028
i: [1,1290,0,4,1,0,0,255] rc6_valued_0
s: 1290
i: [1,1312,0,4,1,0,0,255] rc3_relat_1
s: 1312
d: [2,3019,1,2,1,1,40,1125,2137]
a: [3,3038,1,2,1,3019]
s: 3038
i: [1,1315,0,4,1,0,0,255] rc6_funct_1
s: 1315
d: [2,3024,1,2,1,1,40,1125,2737]
a: [3,3043,1,2,1,3024]
s: 3043
i: [1,1318,0,4,1,0,0,255] rc6_funct_1
s: 1318
i: [1,1319,0,4,1,0,0,255] rc5_finset_1
s: 1319
d: [2,3045,1,2,1,1,40,1127,1969]
a: [3,3051,1,2,1,3045]
s: 3051
i: [1,1322,0,4,1,0,0,255] rc5_finset_1
s: 1322
d: [2,3046,1,2,1,1,40,1127,2142]
a: [3,3056,1,2,1,3046]
s: 3056
i: [1,1323,0,4,1,0,0,255] rc5_finset_1
s: 1323
i: [1,1044,0,4,1,0,0,2] dt_k5_numbers
a: [3,2009,0,4,1,1044]
s: 2009
d: [2,3047,1,2,1,1,40,1127,2417]
a: [3,3061,1,2,1,3047]
s: 3061
i: [1,1045,0,4,1,0,0,1] dt_k4_glib_003
a: [3,2014,0,4,1,1045]
s: 2014
d: [2,3069,1,2,1,1,40,1128,1969]
a: [3,3075,1,2,1,3069]
s: 3075
i: [1,1047,0,4,1,0,0,1] dt_k3_glib_003
a: [3,2020,0,4,1,1047]
s: 2020
i: [1,1048,0,4,1,0,0,1] dt_k2_glib_003
a: [3,2025,0,4,1,1048]
s: 2025
d: [2,3070,1,2,1,1,40,1128,2142]
a: [3,3080,1,2,1,3070]
s: 3080
i: [1,1053,0,4,1,0,0,255] spc7_numerals
a: [3,2052,0,4,1,1053]
s: 2052
d: [2,3071,1,2,1,1,40,1128,2417]
a: [3,3085,1,2,1,3071]
s: 3085
i: [1,1056,0,4,1,0,0,255] spc6_numerals
a: [3,2067,0,4,1,1056]
s: 2067
i: [1,1059,0,4,1,0,0,255] spc5_numerals
a: [3,2082,0,4,1,1059]
s: 2082
d: [2,3089,1,2,1,1,40,1129,1969]
a: [3,3095,1,2,1,3089]
s: 3095
i: [1,1062,0,4,1,0,0,255] spc2_numerals
a: [3,2097,0,4,1,1062]
s: 2097
d: [2,3090,1,2,1,1,40,1129,2142]
a: [3,3100,1,2,1,3090]
s: 3100
i: [1,1065,0,4,1,0,0,255] spc1_numerals
a: [3,2112,0,4,1,1065]
s: 2112
i: [1,1451,0,4,1,0,0,255] rc2_nat_1
a: [3,2747,0,4,1,1451]
s: 2747
d: [2,3091,1,2,1,1,40,1129,2417]
a: [3,3105,1,2,1,3091]
s: 3105
i: [1,1069,0,5,2,0,0,255] rc7_finset_1
s: 1069
d: [2,3107,1,2,1,1,40,1130,1969]
a: [3,3113,1,2,1,3107]
s: 3113
i: [1,1070,0,5,2,0,0,255] rc7_finset_1
s: 1070
i: [1,1072,0,5,2,0,0,2] rc2_finset_1
s: 1072
d: [2,3108,1,2,1,1,40,1130,2142]
a: [3,3118,1,2,1,3108]
s: 3118
i: [1,1073,0,5,2,0,0,2] rc2_finset_1
s: 1073
d: [2,3109,1,2,1,1,40,1130,2417]
a: [3,3123,1,2,1,3109]
s: 3123
i: [1,1077,0,5,2,0,0,2] fc9_card_1
s: 1077
i: [1,1079,0,5,2,0,0,2] rc4_card_1
s: 1079
d: [2,3127,1,2,1,1,40,1131,1969]
a: [3,3133,1,2,1,3127]
s: 3133
i: [1,1104,0,5,2,0,0,255] rc6_card_1
s: 1104
d: [2,3128,1,2,1,1,40,1131,2142]
a: [3,3138,1,2,1,3128]
s: 3138
i: [1,1105,0,5,2,0,0,255] rc6_card_1
s: 1105
i: [1,1117,0,5,2,0,0,255] rc7_finseq_1
s: 1117
d: [2,3129,1,2,1,1,40,1131,2417]
a: [3,3143,1,2,1,3129]
s: 3143
i: [1,1119,0,5,2,0,0,255] rc7_finseq_1
s: 1119
d: [2,3157,1,2,1,1,40,1132,1969]
a: [3,3163,1,2,1,3157]
s: 3163
i: [1,1120,0,5,2,0,0,255] rc7_finseq_1
s: 1120
i: [1,1122,0,5,2,0,0,255] rc7_finseq_1
s: 1122
d: [2,3158,1,2,1,1,40,1132,2142]
a: [3,3168,1,2,1,3158]
s: 3168
i: [1,1123,0,5,2,0,0,255] rc7_finseq_1
s: 1123
d: [2,3159,1,2,1,1,40,1132,2417]
a: [3,3173,1,2,1,3159]
s: 3173
i: [1,1138,0,5,2,0,0,255] t6_boole
s: 1138
i: [1,1139,0,5,2,0,0,255] cc8_card_1
s: 1139
d: [2,3175,1,2,1,1,40,1133,1969]
a: [3,3181,1,2,1,3175]
s: 3181
i: [1,1146,0,5,2,0,0,255] fc31_finset_1
s: 1146
d: [2,3176,1,2,1,1,40,1133,2142]
a: [3,3186,1,2,1,3176]
s: 3186
i: [1,1147,0,5,2,0,0,2] fc17_finset_1
s: 1147
i: [1,1150,0,5,2,0,0,255] cc7_card_1
s: 1150
d: [2,3177,1,2,1,1,40,1133,2417]
a: [3,3191,1,2,1,3177]
s: 3191
i: [1,1151,0,5,2,0,0,255] cc9_card_1
s: 1151
d: [2,3205,1,2,1,1,40,1134,1969]
a: [3,3211,1,2,1,3205]
s: 3211
i: [1,1152,0,5,2,0,0,255] cc9_card_1
s: 1152
i: [1,1153,0,5,2,0,0,255] cc8_ordinal1
s: 1153
d: [2,3206,1,2,1,1,40,1134,2142]
a: [3,3216,1,2,1,3206]
s: 3216
i: [1,1154,0,5,2,0,0,255] cc4_card_1
s: 1154
d: [2,3207,1,2,1,1,40,1134,2417]
a: [3,3221,1,2,1,3207]
s: 3221
i: [1,1155,0,5,2,0,0,255] cc2_nat_1
s: 1155
i: [1,1156,0,5,2,0,0,255] cc25_membered
s: 1156
d: [2,3223,1,2,1,1,40,1135,1969]
a: [3,3229,1,2,1,3223]
s: 3229
i: [1,1269,0,5,2,0,0,255] t7_boole
s: 1269
d: [2,3224,1,2,1,1,40,1135,2142]
a: [3,3234,1,2,1,3224]
s: 3234
i: [1,1313,0,5,1,0,0,255] rc3_relat_1
s: 1313
i: [1,1314,0,5,1,0,0,255] rc3_relat_1
s: 1314
d: [2,3225,1,2,1,1,40,1135,2417]
a: [3,3239,1,2,1,3225]
s: 3239
i: [1,1316,0,5,1,0,0,255] rc6_funct_1
s: 1316
d: [2,3241,1,2,1,1,40,1136,1969]
a: [3,3247,1,2,1,3241]
s: 3247
i: [1,1317,0,5,1,0,0,255] rc6_funct_1
s: 1317
i: [1,1320,0,5,1,0,0,255] rc5_finset_1
s: 1320
d: [2,3242,1,2,1,1,40,1136,2142]
a: [3,3252,1,2,1,3242]
s: 3252
i: [1,1321,0,5,1,0,0,255] rc5_finset_1
s: 1321
d: [2,3243,1,2,1,1,40,1136,2417]
a: [3,3257,1,2,1,3243]
s: 3257
d: [2,3352,0,5,2,1,22,1153,1994]
s: 3352
d: [2,3365,0,5,2,1,22,1154,1994]
s: 3365
d: [2,3259,1,2,1,1,40,1137,1969]
a: [3,3265,1,2,1,3259]
s: 3265
i: [1,1075,0,6,2,0,0,255] rc7_card_1
s: 1075
d: [2,3260,1,2,1,1,40,1137,2142]
a: [3,3270,1,2,1,3260]
s: 3270
i: [1,1103,0,6,2,0,0,255] rc5_card_1
s: 1103
i: [1,1106,0,6,2,0,0,255] rc6_card_1
s: 1106
d: [2,3261,1,2,1,1,40,1137,2417]
a: [3,3275,1,2,1,3261]
s: 3275
i: [1,1118,0,6,2,0,0,255] rc7_finseq_1
s: 1118
d: [2,3378,1,2,1,1,40,1155,2839]
a: [3,3389,1,2,1,3378]
s: 3389
i: [1,1121,0,6,2,0,0,255] rc7_finseq_1
s: 1121
i: [1,1157,0,6,2,0,0,255] cc9_membered
s: 1157
d: [2,3379,1,2,1,1,40,1155,2827]
a: [3,3394,1,2,1,3379]
s: 3394
i: [1,1158,0,6,2,0,0,255] cc10_membered
s: 1158
d: [2,3380,1,2,1,1,40,1155,2815]
a: [3,3399,1,2,1,3380]
s: 3399
i: [1,1159,0,6,2,0,0,255] cc8_membered
s: 1159
i: [1,1160,0,6,2,0,0,255] cc11_membered
s: 1160
d: [2,3381,1,2,1,1,40,1155,2803]
a: [3,3404,1,2,1,3381]
s: 3404
i: [1,1170,0,6,3,0,0,255] cc2_ordinal1
s: 1170
d: [2,3382,1,2,1,1,40,1155,2107]
a: [3,3409,1,2,1,3382]
s: 3409
i: [1,1171,0,6,3,0,0,255] cc6_card_1
s: 1171
i: [1,1178,0,6,3,0,0,255] cc3_relat_1
s: 1178
d: [2,3383,1,2,1,1,40,1155,2092]
a: [3,3414,1,2,1,3383]
s: 3414
d: [2,3710,0,4,2,0,21,1178,1137]
s: 3710
d: [2,3385,1,2,1,1,40,1155,2732]
a: [3,3424,1,2,1,3385]
s: 3424
i: [1,1180,0,6,3,0,0,255] cc9_finset_1
s: 1180
d: [2,3729,0,4,2,0,21,1180,1137]
s: 3729
d: [2,3426,1,2,1,1,40,1156,2839]
a: [3,3437,1,2,1,3426]
s: 3437
i: [1,1182,0,6,3,0,0,255] cc9_valued_0
s: 1182
d: [2,3427,1,2,1,1,40,1156,2827]
a: [3,3442,1,2,1,3427]
s: 3442
d: [2,3754,0,4,2,0,21,1182,1137]
s: 3754
i: [1,1187,0,6,3,0,0,255] cc6_valued_0
s: 1187
d: [2,3428,1,2,1,1,40,1156,2815]
a: [3,3447,1,2,1,3428]
s: 3447
i: [1,1189,0,6,3,0,0,255] cc5_valued_0
s: 1189
d: [2,3430,1,2,1,1,40,1156,2107]
a: [3,3452,1,2,1,3430]
s: 3452
i: [1,1191,0,6,3,0,0,255] cc8_valued_0
s: 1191
i: [1,1199,0,6,3,0,0,255] cc1_finseq_1
s: 1199
d: [2,3431,1,2,1,1,40,1156,2092]
a: [3,3457,1,2,1,3431]
s: 3457
d: [2,3807,0,4,2,0,21,1199,1137]
s: 3807
d: [2,3433,1,2,1,1,40,1156,2732]
a: [3,3467,1,2,1,3433]
s: 3467
i: [1,1201,0,6,3,0,0,255] cc4_finseq_1
s: 1201
i: [1,1265,0,6,2,0,0,255] t1_subset
s: 1265
d: [2,3471,1,2,1,2,40,3352,2839]
a: [3,3482,1,2,1,3471]
s: 3482
i: [1,1294,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 1294
d: [2,3472,1,2,1,2,40,3352,2827]
a: [3,3487,1,2,1,3472]
s: 3487
i: [1,1068,0,7,2,0,0,255] rc7_finset_1
s: 1068
i: [1,1071,0,7,2,0,0,2] rc2_finset_1
s: 1071
d: [2,3473,1,2,1,2,40,3352,2815]
a: [3,3492,1,2,1,3473]
s: 3492
i: [1,1074,0,7,2,0,0,255] rc7_card_1
s: 1074
d: [2,3475,1,2,1,2,40,3352,2107]
a: [3,3497,1,2,1,3475]
s: 3497
i: [1,1078,0,7,2,0,0,2] rc4_card_1
s: 1078
i: [1,1085,0,7,3,0,0,255] cc10_finseq_1
s: 1085
d: [2,3476,1,2,1,2,40,3352,2092]
a: [3,3502,1,2,1,3476]
s: 3502
i: [1,1087,0,7,3,0,0,255] cc7_finset_1
s: 1087
d: [2,3509,1,2,1,2,40,3365,2839]
a: [3,3520,1,2,1,3509]
s: 3520
i: [1,1095,0,7,3,0,0,255] cc8_funct_1
s: 1095
i: [1,1096,0,7,3,0,0,255] cc8_funct_1
s: 1096
d: [2,3510,1,2,1,2,40,3365,2827]
a: [3,3525,1,2,1,3510]
s: 3525
i: [1,1108,0,7,3,0,0,255] cc17_membered
s: 1108
d: [2,3511,1,2,1,2,40,3365,2815]
a: [3,3530,1,2,1,3511]
s: 3530
i: [1,1126,0,7,3,0,0,255] cc5_ordinal1
s: 1126
i: [1,1161,0,7,3,0,0,255] cc10_card_1
s: 1161
d: [2,3513,1,2,1,2,40,3365,2107]
a: [3,3535,1,2,1,3513]
s: 3535
i: [1,1193,0,7,3,0,0,255] cc7_valued_0
s: 1193
d: [2,3514,1,2,1,2,40,3365,2092]
a: [3,3540,1,2,1,3514]
s: 3540
i: [1,1195,0,7,3,0,0,255] cc1_valued_0
s: 1195
i: [1,1236,0,7,3,0,0,255] cc4_valued_0
s: 1236
d: [2,3551,1,5,2,0,40,1103,1152]
a: [3,3565,1,2,1,3551]
s: 3565
i: [1,1239,0,7,3,0,0,255] cc3_valued_0
s: 1239
d: [2,3623,1,2,1,1,40,1159,2009]
a: [3,3629,1,2,1,3623]
s: 3629
i: [1,1244,0,7,3,0,0,255] cc30_valued_0
s: 1244
i: [1,1246,0,7,3,0,0,255] cc31_valued_0
s: 1246
d: [2,3625,1,2,1,1,40,1159,2747]
a: [3,3639,1,2,1,3625]
s: 3639
i: [1,1267,0,7,2,0,0,255] t3_subset
s: 1267
d: [2,3659,1,4,2,1,40,1171,2137]
a: [3,3684,1,2,1,3659]
s: 3684
i: [1,1268,0,7,3,0,0,1] t8_boole
s: 1268
i: [1,1080,0,8,3,0,0,255] cc10_ordinal1
s: 1080
d: [2,3661,1,4,2,1,40,1171,2167]
a: [3,3693,1,2,1,3661]
s: 3693
i: [1,1081,0,8,3,0,0,255] cc20_membered
s: 1081
d: [2,3664,1,4,2,1,40,1171,2397]
d: [2,3695,1,2,1,2,21,3664,2392]
a: [3,3699,1,2,1,3695]
s: 3699
i: [1,1082,0,8,3,0,0,255] cc19_membered
s: 1082
i: [1,1086,0,8,3,0,0,255] cc11_finseq_1
s: 1086
d: [2,3667,1,4,2,1,40,1171,2737]
a: [3,3708,1,2,1,3667]
s: 3708
i: [1,1088,0,8,3,0,0,255] cc8_finset_1
s: 1088
d: [2,3711,1,2,1,1,40,3710,1969]
a: [3,3717,1,2,1,3711]
s: 3717
i: [1,1091,0,8,3,0,0,255] cc23_membered
s: 1091
i: [1,1092,0,8,3,0,0,255] cc22_membered
s: 1092
d: [2,3712,1,2,1,1,40,3710,2142]
a: [3,3722,1,2,1,3712]
s: 3722
i: [1,1097,0,8,3,0,0,255] cc9_funct_1
s: 1097
d: [2,3713,1,2,1,1,40,3710,2417]
a: [3,3727,1,2,1,3713]
s: 3727
i: [1,1098,0,8,3,0,0,255] cc21_membered
s: 1098
i: [1,1109,0,8,3,0,0,255] cc24_membered
s: 1109
d: [2,3730,1,2,1,1,40,3729,1969]
a: [3,3736,1,2,1,3730]
s: 3736
i: [1,1148,0,8,3,0,0,2] cc2_finset_1
s: 1148
d: [2,3731,1,2,1,1,40,3729,2142]
a: [3,3741,1,2,1,3731]
s: 3741
i: [1,1149,0,8,3,0,0,2] cc2_relat_1
s: 1149
i: [1,1169,0,8,4,0,0,255] cc6_funct_1
s: 1169
d: [2,3732,1,2,1,1,40,3729,2417]
a: [3,3746,1,2,1,3732]
s: 3746
i: [1,1176,0,8,4,0,0,255] cc4_funct_1
s: 1176
d: [2,3761,1,2,1,1,40,3754,1969]
a: [3,3767,1,2,1,3761]
s: 3767
d: [2,4696,0,6,3,0,21,1176,1137]
d: [2,4697,0,4,2,0,21,4696,1136]
s: 4697
i: [1,1202,0,8,4,0,0,255] cc5_funct_1
s: 1202
d: [2,3762,1,2,1,1,40,3754,2142]
a: [3,3772,1,2,1,3762]
s: 3772
i: [1,1223,0,8,4,0,0,255] cc6_finseq_1
s: 1223
d: [2,3763,1,2,1,1,40,3754,2417]
a: [3,3777,1,2,1,3763]
s: 3777
i: [1,1226,0,8,4,0,0,255] cc2_finseq_1
s: 1226
i: [1,1233,0,8,4,0,0,255] cc3_funcop_1
s: 1233
d: [2,3791,1,4,2,1,40,1191,2352]
d: [2,3794,1,2,1,2,21,3791,2377]
a: [3,3798,1,2,1,3794]
s: 3798
d: [2,4746,0,6,3,0,21,1233,1137]
d: [2,4747,0,4,2,0,21,4746,1136]
s: 4747
d: [2,3792,1,4,2,1,40,1191,2542]
d: [2,3800,1,2,1,2,21,3792,2552]
a: [3,3804,1,2,1,3800]
s: 3804
i: [1,1241,0,8,3,0,0,255] cc2_valued_0
s: 1241
i: [1,1266,0,8,3,0,0,255] t2_subset
s: 1266
d: [2,3814,1,2,1,1,40,3807,1969]
a: [3,3820,1,2,1,3814]
s: 3820
i: [1,1299,0,8,3,0,0,255] rc6_finset_1
s: 1299
d: [2,3815,1,2,1,1,40,3807,2142]
a: [3,3825,1,2,1,3815]
s: 3825
i: [1,1300,0,8,3,0,0,255] rc6_finset_1
s: 1300
i: [1,1303,0,8,3,0,0,255] rc6_finset_1
s: 1303
d: [2,3816,1,2,1,1,40,3807,2417]
a: [3,3830,1,2,1,3816]
s: 3830
i: [1,1304,0,8,3,0,0,255] rc6_finset_1
s: 1304
d: [2,4015,1,4,2,1,40,1085,2833]
a: [3,4046,1,2,1,4015]
s: 4046
i: [1,1140,0,9,4,0,0,255] cc8_relat_1
s: 1140
i: [1,1143,0,9,4,0,0,1] cc7_relat_1
s: 1143
d: [2,4016,1,4,2,1,40,1085,2839]
a: [3,4055,1,2,1,4016]
s: 4055
i: [1,1218,0,9,4,0,0,255] cc7_finseq_1
s: 1218
d: [2,4112,1,4,2,1,40,1087,2833]
a: [3,4140,1,2,1,4112]
s: 4140
i: [1,1228,0,9,4,0,0,255] cc5_finseq_1
s: 1228
i: [1,1271,0,9,3,0,0,255] t5_subset
s: 1271
d: [2,4163,1,4,2,1,40,1095,2833]
a: [3,4194,1,2,1,4163]
s: 4194
i: [1,1301,0,9,3,0,0,255] rc6_finset_1
s: 1301
d: [2,4164,1,4,2,1,40,1095,2839]
a: [3,4203,1,2,1,4164]
s: 4203
i: [1,1302,0,9,3,0,0,255] rc6_finset_1
s: 1302
i: [1,1165,0,10,5,0,0,255] cc22_valued_0
s: 1165
d: [2,4331,1,4,2,1,40,1108,2692]
a: [3,4346,1,2,1,4331]
s: 4346
i: [1,1166,0,10,5,0,0,255] cc22_valued_0
s: 1166
d: [2,4354,1,4,2,1,40,1126,2833]
a: [3,4376,1,2,1,4354]
s: 4376
i: [1,1184,0,10,4,0,0,255] cc10_valued_0
s: 1184
i: [1,1185,0,10,4,0,0,255] cc11_valued_0
s: 1185
d: [2,4526,1,4,2,1,40,1080,2009]
a: [3,4536,1,2,1,4526]
s: 4536
i: [1,1196,0,10,4,0,0,255] cc12_valued_0
s: 1196
d: [2,4610,1,4,2,1,40,1091,2747]
a: [3,4614,1,2,1,4610]
s: 4614
i: [1,1197,0,10,4,0,0,255] cc15_valued_0
s: 1197
i: [1,1234,0,10,4,0,0,2] cc3_funct_1
s: 1234
d: [2,4621,1,4,2,1,40,1092,2009]
a: [3,4631,1,2,1,4621]
s: 4631
i: [1,1270,0,10,3,0,0,255] t4_subset
s: 1270
d: [2,4683,1,4,2,1,40,1149,2009]
a: [3,4693,1,2,1,4683]
s: 4693
i: [1,1305,0,10,4,0,0,255] fc19_funct_1
s: 1305
r: 3684
r: 3708
r: 4536
r: 4614
i: [1,1306,0,10,4,0,0,255] fc35_finset_1
s: 1306
a: [3,3680,1,2,1,3659]
s: 3680
r: 3680
i: [1,1307,0,10,4,0,0,255] fc8_funcop_1
s: 1307
a: [3,3679,1,2,1,3659]
s: 3679
i: [1,1308,0,10,4,0,0,255] fc8_funcop_1
s: 1308
i: [1,1309,0,10,4,0,0,255] fc66_valued_0
s: 1309
s: 3684
i: [1,1101,0,11,5,0,0,255] cc1_funcop_1
s: 1101
a: [3,3704,1,2,1,3667]
s: 3704
r: 3704
i: [1,1237,0,12,4,0,0,255] cc13_valued_0
s: 1237
i: [1,1242,0,12,4,0,0,255] cc14_valued_0
s: 1242
a: [3,3703,1,2,1,3667]
s: 3703
i: [1,1310,0,12,4,0,0,255] cc6_relat_1
s: 1310
r: 1242
r: 1237
s: 3708
i: [1,1311,0,12,4,0,0,1] cc5_relat_1
s: 1311
i: [1,1324,0,13,5,0,0,255] fc29_finseq_1
s: 1324
d: [2,3868,1,5,2,0,40,1068,1157]
a: [3,3895,1,2,1,3868]
s: 3895
d: [2,5246,2,0,0,2,52,1929,3895,1076]
e: 5246
d: [2,5250,2,2,1,1,40,3895,1076]
d: [2,5252,2,0,0,2,21,5250,1929]
e: 5252
r: 3895
i: [1,1325,0,13,5,0,0,255] fc17_funct_1
s: 1325
a: [3,3894,1,2,1,3868]
s: 3894
i: [1,1326,0,13,5,0,0,255] fc17_funct_1
s: 1326
i: [1,1298,0,14,4,0,0,2] existence_m2_subset_1
s: 1298
d: [2,3869,1,5,2,0,40,1068,1159]
a: [3,3904,1,2,1,3869]
s: 3904
d: [2,5254,2,0,0,2,52,1939,3904,1076]
e: 5254
d: [2,5258,2,2,1,1,40,3904,1076]
d: [2,5260,2,0,0,2,21,5258,1939]
e: 5260
r: 3904
i: [1,1264,0,15,5,0,0,255] d7_glib_000
s: 1264
a: [3,3903,1,2,1,3869]
s: 3903
i: [1,1295,0,15,5,0,0,2] dt_m2_subset_1
s: 1295
i: [1,1296,0,15,5,0,0,2] redefinition_m2_subset_1
s: 1296
d: [2,4014,1,5,2,1,40,1085,2009]
a: [3,4041,1,2,1,4014]
s: 4041
r: 4041
i: [1,1297,0,15,5,0,0,2] redefinition_m2_subset_1
s: 1297
a: [3,4040,1,2,1,4014]
s: 4040
d: [2,5262,0,15,5,1,22,1264,1989]
s: 5262
i: [1,1250,0,16,7,0,0,1] dt_k6_glib_003
s: 1250
d: [2,5319,1,13,6,1,40,1250,1490]
d: [2,5337,1,11,5,2,21,5319,1475]
d: [2,5338,1,9,4,3,21,5337,1480]
d: [2,5339,1,7,3,4,21,5338,1485]
d: [2,5340,1,5,2,5,21,5339,1495]
a: [3,5348,1,2,1,5340]
s: 5348
i: [1,1251,0,16,7,0,0,1] dt_k6_glib_003
s: 1251
d: [2,4033,1,5,2,1,40,1085,2747]
a: [3,4095,1,2,1,4033]
s: 4095
r: 4095
i: [1,1252,0,16,7,0,0,1] dt_k6_glib_003
s: 1252
i: [1,1253,0,16,7,0,0,1] dt_k6_glib_003
s: 1253
a: [3,4094,1,2,1,4033]
s: 4094
i: [1,1254,0,16,7,0,0,1] dt_k7_glib_003
s: 1254
d: [2,5599,1,13,6,1,40,1254,1490]
d: [2,5617,1,11,5,2,21,5599,1475]
d: [2,5618,1,9,4,3,21,5617,1480]
d: [2,5619,1,7,3,4,21,5618,1485]
d: [2,5620,1,5,2,5,21,5619,1495]
a: [3,5628,1,2,1,5620]
s: 5628
i: [1,1255,0,16,7,0,0,1] dt_k7_glib_003
s: 1255
i: [1,1256,0,16,7,0,0,1] dt_k7_glib_003
s: 1256
d: [2,4162,1,5,2,1,40,1095,2009]
a: [3,4189,1,2,1,4162]
s: 4189
i: [1,1257,0,16,7,0,0,1] dt_k7_glib_003
s: 1257
d: [2,4181,1,5,2,1,40,1095,2747]
a: [3,4242,1,2,1,4181]
s: 4242
i: [1,1258,0,16,7,0,0,255] dt_k5_glib_003
s: 1258
r: 1495
a: [3,1494,0,2,1,1012]
s: 1494
i: [1,969,0,14,3,1,0,0] t4_glib_003
a: [3,1566,0,2,1,969]
s: 1566
i: [1,920,0,14,3,1,0,0] t4_glib_003
a: [3,1619,0,2,1,920]
s: 1619
i: [1,871,0,14,3,1,0,0] t4_glib_003
a: [3,1672,0,2,1,871]
s: 1672
i: [1,822,0,14,3,1,0,0] t4_glib_003
a: [3,1725,0,2,1,822]
s: 1725
i: [1,724,0,14,3,1,0,0] t4_glib_003
a: [3,1831,0,2,1,724]
s: 1831
i: [1,773,0,15,3,1,0,0] t4_glib_003
a: [3,1778,0,3,1,773]
s: 1778
a: [3,1465,0,6,1,1018]
s: 1465
d: [2,5968,1,13,6,1,40,1778,1258]
d: [2,5974,1,11,5,2,21,5968,1619]
d: [2,5975,1,9,4,3,21,5974,1672]
d: [2,5976,1,7,3,4,21,5975,1725]
d: [2,5977,1,5,2,5,21,5976,1831]
a: [3,5985,1,2,1,5977]
s: 5985
i: [1,1260,0,16,7,0,0,255] dt_k5_glib_003
s: 1260
i: [1,1261,0,16,7,0,0,255] dt_k5_glib_003
s: 1261
d: [2,5971,1,13,6,1,40,1778,1251]
d: [2,6007,1,11,5,2,21,5971,1619]
d: [2,6008,1,9,4,3,21,6007,1672]
d: [2,6009,1,7,3,4,21,6008,1725]
d: [2,6010,1,5,2,5,21,6009,1831]
a: [3,6018,1,2,1,6010]
s: 6018
i: [1,1263,0,16,7,0,0,255] dt_k5_glib_003
s: 1263
d: [2,5969,1,13,6,1,40,1778,1255]
d: [2,5987,1,11,5,2,21,5969,1566]
d: [2,5988,1,9,4,3,21,5987,1619]
d: [2,5989,1,7,3,4,21,5988,1672]
d: [2,5990,1,5,2,5,21,5989,1725]
d: [2,5991,1,3,1,6,21,5990,1831]
a: [3,5995,1,3,1,5991]
s: 5995
i: [1,1259,0,18,7,0,0,255] dt_k5_glib_003
s: 1259
i: [1,1262,0,18,7,0,0,255] dt_k5_glib_003
s: 1262
d: [2,5970,1,13,6,1,40,1778,1254]
d: [2,5997,1,11,5,2,21,5970,1566]
d: [2,5998,1,9,4,3,21,5997,1619]
d: [2,5999,1,7,3,4,21,5998,1672]
d: [2,6000,1,5,2,5,21,5999,1725]
d: [2,6001,1,3,1,6,21,6000,1831]
a: [3,6005,1,3,1,6001]
s: 6005
i: [1,1247,0,19,7,0,0,1] d8_glib_003
s: 1247
d: [2,4244,1,5,2,1,40,1096,2009]
a: [3,4267,1,2,1,4244]
s: 4267
i: [1,1248,0,19,7,0,0,1] d9_glib_003
s: 1248
d: [2,6145,1,0,0,7,52,1831,1566,1725,1672,1619,1778,1465,1248]
e: 6145
d: [2,6154,1,18,7,1,44,1465,1248]
d: [2,6155,1,13,6,1,19,6154]
d: [2,6165,1,11,5,2,21,6155,1566]
d: [2,6166,1,9,4,3,21,6165,1619]
d: [2,6167,1,7,3,4,21,6166,1672]
d: [2,6168,1,5,2,5,21,6167,1725]
d: [2,6169,1,2,1,6,21,6168,1778]
d: [2,6170,1,0,0,7,21,6169,1831]
e: 6170
r: 1490
r: 1494
r: 1831
s: 1495
i: [1,1011,0,14,3,1,0,0] t4_glib_003
a: [3,1500,0,2,1,1011]
s: 1500
i: [1,1004,0,14,3,1,0,0] t4_glib_003
a: [3,1511,0,2,1,1004]
s: 1511
i: [1,997,0,14,3,1,0,0] t4_glib_003
a: [3,1522,0,2,1,997]
s: 1522
i: [1,990,0,14,3,1,0,0] t4_glib_003
a: [3,1533,0,2,1,990]
s: 1533
i: [1,976,0,14,3,1,0,0] t4_glib_003
a: [3,1555,0,2,1,976]
s: 1555
a: [3,1830,0,2,1,724]
s: 1830
a: [3,1489,0,3,1,1013]
s: 1489
d: [2,6173,1,6,3,1,40,1489,1228]
d: [2,6175,1,4,2,2,21,6173,1485]
d: [2,6176,1,2,1,3,21,6175,1495]
a: [3,6180,1,2,1,6176]
s: 6180
i: [1,983,0,15,3,1,0,0] t4_glib_003
a: [3,1544,0,3,1,983]
s: 1544
a: [3,1461,0,6,1,1018]
s: 1461
d: [2,6260,1,0,0,7,52,1555,1500,1533,1522,1511,1544,1461,1247]
e: 6260
d: [2,6261,1,18,7,1,44,1461,1247]
d: [2,6262,1,13,6,1,19,6261]
d: [2,6264,1,11,5,2,21,6262,1500]
d: [2,6265,1,9,4,3,21,6264,1511]
d: [2,6266,1,7,3,4,21,6265,1522]
d: [2,6267,1,5,2,5,21,6266,1533]
d: [2,6268,1,2,1,6,21,6267,1544]
d: [2,6269,1,0,0,7,21,6268,1555]
e: 6269
r: 1489
r: 1555
a: [3,1554,0,2,1,976]
s: 1554
s: 1490
d: [2,6174,1,6,3,1,40,1489,1218]
d: [2,6182,1,4,2,2,21,6174,1485]
d: [2,6183,1,2,1,3,21,6182,1495]
a: [3,6187,1,2,1,6183]
s: 6187
i: [1,1249,0,19,7,0,0,1] d7_glib_003
s: 1249
d: [2,6301,1,0,0,7,52,1495,1470,1485,1480,1475,1490,1457,1249]
e: 6301
d: [2,6310,1,18,7,1,44,1457,1249]
d: [2,6311,1,13,6,1,19,6310]
d: [2,6321,1,11,5,2,21,6311,1470]
d: [2,6322,1,9,4,3,21,6321,1475]
d: [2,6323,1,7,3,4,21,6322,1480]
d: [2,6324,1,5,2,5,21,6323,1485]
d: [2,6325,1,2,1,6,21,6324,1490]
d: [2,6326,1,0,0,7,21,6325,1495]
e: 6326
i: [1,1010,0,10,3,1,0,0] t4_glib_003
i: [1,1009,0,10,3,1,0,0] t4_glib_003
i: [1,1008,0,10,3,1,0,0] t4_glib_003
i: [1,1007,0,10,3,1,0,0] t4_glib_003
i: [1,1006,0,11,3,1,0,0] t4_glib_003
i: [1,1005,0,10,3,1,0,0] t4_glib_003
i: [1,1003,0,10,3,1,0,0] t4_glib_003
i: [1,1002,0,10,3,1,0,0] t4_glib_003
i: [1,1001,0,10,3,1,0,0] t4_glib_003
i: [1,1000,0,10,3,1,0,0] t4_glib_003
i: [1,999,0,11,3,1,0,0] t4_glib_003
i: [1,998,0,10,3,1,0,0] t4_glib_003
i: [1,996,0,10,3,1,0,0] t4_glib_003
i: [1,995,0,10,3,1,0,0] t4_glib_003
i: [1,994,0,10,3,1,0,0] t4_glib_003
i: [1,993,0,10,3,1,0,0] t4_glib_003
i: [1,992,0,11,3,1,0,0] t4_glib_003
i: [1,991,0,10,3,1,0,0] t4_glib_003
i: [1,989,0,10,3,1,0,0] t4_glib_003
i: [1,988,0,10,3,1,0,0] t4_glib_003
i: [1,987,0,10,3,1,0,0] t4_glib_003
i: [1,986,0,10,3,1,0,0] t4_glib_003
i: [1,985,0,11,3,1,0,0] t4_glib_003
i: [1,984,0,10,3,1,0,0] t4_glib_003
i: [1,982,0,11,3,1,0,0] t4_glib_003
i: [1,981,0,11,3,1,0,0] t4_glib_003
i: [1,980,0,11,3,1,0,0] t4_glib_003
i: [1,979,0,11,3,1,0,0] t4_glib_003
i: [1,978,0,12,3,1,0,0] t4_glib_003
i: [1,977,0,11,3,1,0,0] t4_glib_003
i: [1,975,0,10,3,1,0,0] t4_glib_003
i: [1,974,0,10,3,1,0,0] t4_glib_003
i: [1,973,0,10,3,1,0,0] t4_glib_003
i: [1,972,0,10,3,1,0,0] t4_glib_003
i: [1,971,0,11,3,1,0,0] t4_glib_003
i: [1,970,0,10,3,1,0,0] t4_glib_003
i: [1,968,0,10,3,1,0,0] t4_glib_003
i: [1,967,0,10,3,1,0,0] t4_glib_003
i: [1,966,0,10,3,1,0,0] t4_glib_003
i: [1,965,0,10,3,1,0,0] t4_glib_003
i: [1,964,0,11,3,1,0,0] t4_glib_003
i: [1,963,0,10,3,1,0,0] t4_glib_003
i: [1,962,0,10,3,1,0,0] t4_glib_003
i: [1,961,0,6,3,1,0,0] t4_glib_003
i: [1,960,0,6,3,1,0,0] t4_glib_003
i: [1,959,0,6,3,1,0,0] t4_glib_003
i: [1,958,0,6,3,1,0,0] t4_glib_003
i: [1,957,0,7,3,1,0,0] t4_glib_003
i: [1,956,0,6,3,1,0,0] t4_glib_003
i: [1,955,0,10,3,1,0,0] t4_glib_003
i: [1,954,0,6,3,1,0,0] t4_glib_003
i: [1,953,0,6,3,1,0,0] t4_glib_003
i: [1,952,0,6,3,1,0,0] t4_glib_003
i: [1,951,0,6,3,1,0,0] t4_glib_003
i: [1,950,0,7,3,1,0,0] t4_glib_003
i: [1,949,0,6,3,1,0,0] t4_glib_003
i: [1,948,0,10,3,1,0,0] t4_glib_003
i: [1,947,0,6,3,1,0,0] t4_glib_003
i: [1,946,0,6,3,1,0,0] t4_glib_003
i: [1,945,0,6,3,1,0,0] t4_glib_003
i: [1,944,0,6,3,1,0,0] t4_glib_003
i: [1,943,0,7,3,1,0,0] t4_glib_003
i: [1,942,0,6,3,1,0,0] t4_glib_003
i: [1,941,0,10,3,1,0,0] t4_glib_003
i: [1,940,0,6,3,1,0,0] t4_glib_003
i: [1,939,0,6,3,1,0,0] t4_glib_003
i: [1,938,0,6,3,1,0,0] t4_glib_003
i: [1,937,0,6,3,1,0,0] t4_glib_003
i: [1,936,0,7,3,1,0,0] t4_glib_003
i: [1,935,0,6,3,1,0,0] t4_glib_003
i: [1,934,0,11,3,1,0,0] t4_glib_003
i: [1,933,0,7,3,1,0,0] t4_glib_003
i: [1,932,0,7,3,1,0,0] t4_glib_003
i: [1,931,0,7,3,1,0,0] t4_glib_003
i: [1,930,0,7,3,1,0,0] t4_glib_003
i: [1,929,0,8,3,1,0,0] t4_glib_003
i: [1,928,0,7,3,1,0,0] t4_glib_003
i: [1,927,0,10,3,1,0,0] t4_glib_003
i: [1,926,0,6,3,1,0,0] t4_glib_003
i: [1,925,0,6,3,1,0,0] t4_glib_003
i: [1,924,0,6,3,1,0,0] t4_glib_003
i: [1,923,0,6,3,1,0,0] t4_glib_003
i: [1,922,0,7,3,1,0,0] t4_glib_003
i: [1,921,0,6,3,1,0,0] t4_glib_003
i: [1,919,0,10,3,1,0,0] t4_glib_003
i: [1,918,0,10,3,1,0,0] t4_glib_003
i: [1,917,0,10,3,1,0,0] t4_glib_003
i: [1,916,0,10,3,1,0,0] t4_glib_003
i: [1,915,0,11,3,1,0,0] t4_glib_003
i: [1,914,0,10,3,1,0,0] t4_glib_003
i: [1,913,0,10,3,1,0,0] t4_glib_003
i: [1,912,0,6,3,1,0,0] t4_glib_003
i: [1,911,0,6,3,1,0,0] t4_glib_003
i: [1,910,0,6,3,1,0,0] t4_glib_003
i: [1,909,0,6,3,1,0,0] t4_glib_003
i: [1,908,0,7,3,1,0,0] t4_glib_003
i: [1,907,0,6,3,1,0,0] t4_glib_003
i: [1,906,0,10,3,1,0,0] t4_glib_003
i: [1,905,0,6,3,1,0,0] t4_glib_003
i: [1,904,0,6,3,1,0,0] t4_glib_003
i: [1,903,0,6,3,1,0,0] t4_glib_003
i: [1,902,0,6,3,1,0,0] t4_glib_003
i: [1,901,0,7,3,1,0,0] t4_glib_003
i: [1,900,0,6,3,1,0,0] t4_glib_003
i: [1,899,0,10,3,1,0,0] t4_glib_003
i: [1,898,0,6,3,1,0,0] t4_glib_003
i: [1,897,0,6,3,1,0,0] t4_glib_003
i: [1,896,0,6,3,1,0,0] t4_glib_003
i: [1,895,0,6,3,1,0,0] t4_glib_003
i: [1,894,0,7,3,1,0,0] t4_glib_003
i: [1,893,0,6,3,1,0,0] t4_glib_003
i: [1,892,0,10,3,1,0,0] t4_glib_003
i: [1,891,0,6,3,1,0,0] t4_glib_003
i: [1,890,0,6,3,1,0,0] t4_glib_003
i: [1,889,0,6,3,1,0,0] t4_glib_003
i: [1,888,0,6,3,1,0,0] t4_glib_003
i: [1,887,0,7,3,1,0,0] t4_glib_003
i: [1,886,0,6,3,1,0,0] t4_glib_003
i: [1,885,0,11,3,1,0,0] t4_glib_003
i: [1,884,0,7,3,1,0,0] t4_glib_003
i: [1,883,0,7,3,1,0,0] t4_glib_003
i: [1,882,0,7,3,1,0,0] t4_glib_003
i: [1,881,0,7,3,1,0,0] t4_glib_003
i: [1,880,0,8,3,1,0,0] t4_glib_003
i: [1,879,0,7,3,1,0,0] t4_glib_003
i: [1,878,0,10,3,1,0,0] t4_glib_003
i: [1,877,0,6,3,1,0,0] t4_glib_003
i: [1,876,0,6,3,1,0,0] t4_glib_003
i: [1,875,0,6,3,1,0,0] t4_glib_003
i: [1,874,0,6,3,1,0,0] t4_glib_003
i: [1,873,0,7,3,1,0,0] t4_glib_003
i: [1,872,0,6,3,1,0,0] t4_glib_003
i: [1,870,0,10,3,1,0,0] t4_glib_003
i: [1,869,0,10,3,1,0,0] t4_glib_003
i: [1,868,0,10,3,1,0,0] t4_glib_003
i: [1,867,0,10,3,1,0,0] t4_glib_003
i: [1,866,0,11,3,1,0,0] t4_glib_003
i: [1,865,0,10,3,1,0,0] t4_glib_003
i: [1,864,0,10,3,1,0,0] t4_glib_003
i: [1,863,0,6,3,1,0,0] t4_glib_003
i: [1,862,0,6,3,1,0,0] t4_glib_003
i: [1,861,0,6,3,1,0,0] t4_glib_003
i: [1,860,0,6,3,1,0,0] t4_glib_003
i: [1,859,0,7,3,1,0,0] t4_glib_003
i: [1,858,0,6,3,1,0,0] t4_glib_003
i: [1,857,0,10,3,1,0,0] t4_glib_003
i: [1,856,0,6,3,1,0,0] t4_glib_003
i: [1,855,0,6,3,1,0,0] t4_glib_003
i: [1,854,0,6,3,1,0,0] t4_glib_003
i: [1,853,0,6,3,1,0,0] t4_glib_003
i: [1,852,0,7,3,1,0,0] t4_glib_003
i: [1,851,0,6,3,1,0,0] t4_glib_003
i: [1,850,0,10,3,1,0,0] t4_glib_003
i: [1,849,0,6,3,1,0,0] t4_glib_003
i: [1,848,0,6,3,1,0,0] t4_glib_003
i: [1,847,0,6,3,1,0,0] t4_glib_003
i: [1,846,0,6,3,1,0,0] t4_glib_003
i: [1,845,0,7,3,1,0,0] t4_glib_003
i: [1,844,0,6,3,1,0,0] t4_glib_003
i: [1,843,0,10,3,1,0,0] t4_glib_003
i: [1,842,0,6,3,1,0,0] t4_glib_003
i: [1,841,0,6,3,1,0,0] t4_glib_003
i: [1,840,0,6,3,1,0,0] t4_glib_003
i: [1,839,0,6,3,1,0,0] t4_glib_003
i: [1,838,0,7,3,1,0,0] t4_glib_003
i: [1,837,0,6,3,1,0,0] t4_glib_003
i: [1,836,0,11,3,1,0,0] t4_glib_003
i: [1,835,0,7,3,1,0,0] t4_glib_003
i: [1,834,0,7,3,1,0,0] t4_glib_003
i: [1,833,0,7,3,1,0,0] t4_glib_003
i: [1,832,0,7,3,1,0,0] t4_glib_003
i: [1,831,0,8,3,1,0,0] t4_glib_003
i: [1,830,0,7,3,1,0,0] t4_glib_003
i: [1,829,0,10,3,1,0,0] t4_glib_003
i: [1,828,0,6,3,1,0,0] t4_glib_003
i: [1,827,0,6,3,1,0,0] t4_glib_003
i: [1,826,0,6,3,1,0,0] t4_glib_003
i: [1,825,0,6,3,1,0,0] t4_glib_003
i: [1,824,0,7,3,1,0,0] t4_glib_003
i: [1,823,0,6,3,1,0,0] t4_glib_003
i: [1,821,0,10,3,1,0,0] t4_glib_003
i: [1,820,0,10,3,1,0,0] t4_glib_003
i: [1,819,0,10,3,1,0,0] t4_glib_003
i: [1,818,0,10,3,1,0,0] t4_glib_003
i: [1,817,0,11,3,1,0,0] t4_glib_003
i: [1,816,0,10,3,1,0,0] t4_glib_003
i: [1,815,0,10,3,1,0,0] t4_glib_003
i: [1,814,0,6,3,1,0,0] t4_glib_003
i: [1,813,0,6,3,1,0,0] t4_glib_003
i: [1,812,0,6,3,1,0,0] t4_glib_003
i: [1,811,0,6,3,1,0,0] t4_glib_003
i: [1,810,0,7,3,1,0,0] t4_glib_003
i: [1,809,0,6,3,1,0,0] t4_glib_003
i: [1,808,0,10,3,1,0,0] t4_glib_003
i: [1,807,0,6,3,1,0,0] t4_glib_003
i: [1,806,0,6,3,1,0,0] t4_glib_003
i: [1,805,0,6,3,1,0,0] t4_glib_003
i: [1,804,0,6,3,1,0,0] t4_glib_003
i: [1,803,0,7,3,1,0,0] t4_glib_003
i: [1,802,0,6,3,1,0,0] t4_glib_003
i: [1,801,0,10,3,1,0,0] t4_glib_003
i: [1,800,0,6,3,1,0,0] t4_glib_003
i: [1,799,0,6,3,1,0,0] t4_glib_003
i: [1,798,0,6,3,1,0,0] t4_glib_003
i: [1,797,0,6,3,1,0,0] t4_glib_003
i: [1,796,0,7,3,1,0,0] t4_glib_003
i: [1,795,0,6,3,1,0,0] t4_glib_003
i: [1,794,0,10,3,1,0,0] t4_glib_003
i: [1,793,0,6,3,1,0,0] t4_glib_003
i: [1,792,0,6,3,1,0,0] t4_glib_003
i: [1,791,0,6,3,1,0,0] t4_glib_003
i: [1,790,0,6,3,1,0,0] t4_glib_003
i: [1,789,0,7,3,1,0,0] t4_glib_003
i: [1,788,0,6,3,1,0,0] t4_glib_003
i: [1,787,0,11,3,1,0,0] t4_glib_003
i: [1,786,0,7,3,1,0,0] t4_glib_003
i: [1,785,0,7,3,1,0,0] t4_glib_003
i: [1,784,0,7,3,1,0,0] t4_glib_003
i: [1,783,0,7,3,1,0,0] t4_glib_003
i: [1,782,0,8,3,1,0,0] t4_glib_003
i: [1,781,0,7,3,1,0,0] t4_glib_003
i: [1,780,0,10,3,1,0,0] t4_glib_003
i: [1,779,0,6,3,1,0,0] t4_glib_003
i: [1,778,0,6,3,1,0,0] t4_glib_003
i: [1,777,0,6,3,1,0,0] t4_glib_003
i: [1,776,0,6,3,1,0,0] t4_glib_003
i: [1,775,0,7,3,1,0,0] t4_glib_003
i: [1,774,0,6,3,1,0,0] t4_glib_003
i: [1,772,0,11,3,1,0,0] t4_glib_003
i: [1,771,0,11,3,1,0,0] t4_glib_003
i: [1,770,0,11,3,1,0,0] t4_glib_003
i: [1,769,0,11,3,1,0,0] t4_glib_003
i: [1,768,0,12,3,1,0,0] t4_glib_003
i: [1,767,0,11,3,1,0,0] t4_glib_003
i: [1,766,0,11,3,1,0,0] t4_glib_003
i: [1,765,0,7,3,1,0,0] t4_glib_003
i: [1,764,0,7,3,1,0,0] t4_glib_003
i: [1,763,0,7,3,1,0,0] t4_glib_003
i: [1,762,0,7,3,1,0,0] t4_glib_003
i: [1,761,0,8,3,1,0,0] t4_glib_003
i: [1,760,0,7,3,1,0,0] t4_glib_003
i: [1,759,0,11,3,1,0,0] t4_glib_003
i: [1,758,0,7,3,1,0,0] t4_glib_003
i: [1,757,0,7,3,1,0,0] t4_glib_003
i: [1,756,0,7,3,1,0,0] t4_glib_003
i: [1,755,0,7,3,1,0,0] t4_glib_003
i: [1,754,0,8,3,1,0,0] t4_glib_003
i: [1,753,0,7,3,1,0,0] t4_glib_003
i: [1,752,0,11,3,1,0,0] t4_glib_003
i: [1,751,0,7,3,1,0,0] t4_glib_003
i: [1,750,0,7,3,1,0,0] t4_glib_003
i: [1,749,0,7,3,1,0,0] t4_glib_003
i: [1,748,0,7,3,1,0,0] t4_glib_003
i: [1,747,0,8,3,1,0,0] t4_glib_003
i: [1,746,0,7,3,1,0,0] t4_glib_003
i: [1,745,0,11,3,1,0,0] t4_glib_003
i: [1,744,0,7,3,1,0,0] t4_glib_003
i: [1,743,0,7,3,1,0,0] t4_glib_003
i: [1,742,0,7,3,1,0,0] t4_glib_003
i: [1,741,0,7,3,1,0,0] t4_glib_003
i: [1,740,0,8,3,1,0,0] t4_glib_003
i: [1,739,0,7,3,1,0,0] t4_glib_003
i: [1,738,0,12,3,1,0,0] t4_glib_003
i: [1,737,0,8,3,1,0,0] t4_glib_003
i: [1,736,0,8,3,1,0,0] t4_glib_003
i: [1,735,0,8,3,1,0,0] t4_glib_003
i: [1,734,0,8,3,1,0,0] t4_glib_003
i: [1,733,0,9,3,1,0,0] t4_glib_003
i: [1,732,0,8,3,1,0,0] t4_glib_003
i: [1,731,0,11,3,1,0,0] t4_glib_003
i: [1,730,0,7,3,1,0,0] t4_glib_003
i: [1,729,0,7,3,1,0,0] t4_glib_003
i: [1,728,0,7,3,1,0,0] t4_glib_003
i: [1,727,0,7,3,1,0,0] t4_glib_003
i: [1,726,0,8,3,1,0,0] t4_glib_003
i: [1,725,0,7,3,1,0,0] t4_glib_003
i: [1,723,0,10,3,1,0,0] t4_glib_003
i: [1,722,0,10,3,1,0,0] t4_glib_003
i: [1,721,0,10,3,1,0,0] t4_glib_003
i: [1,720,0,10,3,1,0,0] t4_glib_003
i: [1,719,0,11,3,1,0,0] t4_glib_003
i: [1,718,0,10,3,1,0,0] t4_glib_003
i: [1,717,0,10,3,1,0,0] t4_glib_003
i: [1,716,0,6,3,1,0,0] t4_glib_003
i: [1,715,0,6,3,1,0,0] t4_glib_003
i: [1,714,0,6,3,1,0,0] t4_glib_003
i: [1,713,0,6,3,1,0,0] t4_glib_003
i: [1,712,0,7,3,1,0,0] t4_glib_003
i: [1,711,0,6,3,1,0,0] t4_glib_003
i: [1,710,0,10,3,1,0,0] t4_glib_003
i: [1,709,0,6,3,1,0,0] t4_glib_003
i: [1,708,0,6,3,1,0,0] t4_glib_003
i: [1,707,0,6,3,1,0,0] t4_glib_003
i: [1,706,0,6,3,1,0,0] t4_glib_003
i: [1,705,0,7,3,1,0,0] t4_glib_003
i: [1,704,0,6,3,1,0,0] t4_glib_003
i: [1,703,0,10,3,1,0,0] t4_glib_003
i: [1,702,0,6,3,1,0,0] t4_glib_003
i: [1,701,0,6,3,1,0,0] t4_glib_003
i: [1,700,0,6,3,1,0,0] t4_glib_003
i: [1,699,0,6,3,1,0,0] t4_glib_003
i: [1,698,0,7,3,1,0,0] t4_glib_003
i: [1,697,0,6,3,1,0,0] t4_glib_003
i: [1,696,0,10,3,1,0,0] t4_glib_003
i: [1,695,0,6,3,1,0,0] t4_glib_003
i: [1,694,0,6,3,1,0,0] t4_glib_003
i: [1,693,0,6,3,1,0,0] t4_glib_003
i: [1,692,0,6,3,1,0,0] t4_glib_003
i: [1,691,0,7,3,1,0,0] t4_glib_003
i: [1,690,0,6,3,1,0,0] t4_glib_003
i: [1,689,0,11,3,1,0,0] t4_glib_003
i: [1,688,0,7,3,1,0,0] t4_glib_003
i: [1,687,0,7,3,1,0,0] t4_glib_003
i: [1,686,0,7,3,1,0,0] t4_glib_003
i: [1,685,0,7,3,1,0,0] t4_glib_003
i: [1,684,0,8,3,1,0,0] t4_glib_003
i: [1,683,0,7,3,1,0,0] t4_glib_003
i: [1,682,0,10,3,1,0,0] t4_glib_003
i: [1,681,0,6,3,1,0,0] t4_glib_003
i: [1,680,0,6,3,1,0,0] t4_glib_003
i: [1,679,0,6,3,1,0,0] t4_glib_003
i: [1,678,0,6,3,1,0,0] t4_glib_003
i: [1,677,0,7,3,1,0,0] t4_glib_003
i: [1,676,0,6,3,1,0,0] t4_glib_003
f: 6326,6268,6169,676,677,678,679,680,681,682,683,684,685,686,687,688,689,690,691,692,693,694,695,696,697,698,699,700,701,702,703,704,705,706,707,708,709,710,711,712,713,714,715,716,717,718,719,720,721,722,723,724,725,726,727,728,729,730,731,732,733,734,735,736,737,738,739,740,741,742,743,744,745,746,747,748,749,750,751,752,753,754,755,756,757,758,759,760,761,762,763,764,765,766,767,768,769,770,771,772,773,774,775,776,777,778,779,780,781,782,783,784,785,786,787,788,789,790,791,792,793,794,795,796,797,798,799,800,801,802,803,804,805,806,807,808,809,810,811,812,813,814,815,816,817,818,819,820,821,822,823,824,825,826,827,828,829,830,831,832,833,834,835,836,837,838,839,840,841,842,843,844,845,846,847,848,849,850,851,852,853,854,855,856,857,858,859,860,861,862,863,864,865,866,867,868,869,870,871,872,873,874,875,876,877,878,879,880,881,882,883,884,885,886,887,888,889,890,891,892,893,894,895,896,897,898,899,900,901,902,903,904,905,906,907,908,909,910,911,912,913,914,915,916,917,918,919,920,921,922,923,924,925,926,927,928,929,930,931,932,933,934,935,936,937,938,939,940,941,942,943,944,945,946,947,948,949,950,951,952,953,954,955,956,957,958,959,960,961,962,963,964,965,966,967,968,969,970,971,972,973,974,975,976,977,978,979,980,981,982,983,984,985,986,987,988,989,990,991,992,993,994,995,996,997,998,999,1000,1001,1002,1003,1004,1005,1006,1007,1008,1009,1010,1011,1012,1013,1014,1015,1016,1017,1018
% Refutation found. Thanks to Tanya!
% SZS status Theorem for glib_003__t4_glib_003
% SZS output start Proof for glib_003__t4_glib_003
1. ! [X0] : ((v3_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k7_glib_003(X0) = k1_funct_1(X0,k4_glib_003)) & ! [X0] : ((v2_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k6_glib_003(X0) = k1_funct_1(X0,k3_glib_003)) & ! [X0] : ((v1_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k5_glib_003(X0) = k1_funct_1(X0,k2_glib_003)) [input t4_glib_003]
2. ~(! [X0] : ((v3_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k7_glib_003(X0) = k1_funct_1(X0,k4_glib_003)) & ! [X0] : ((v2_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k6_glib_003(X0) = k1_funct_1(X0,k3_glib_003)) & ! [X0] : ((v1_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k5_glib_003(X0) = k1_funct_1(X0,k2_glib_003))) [negated conjecture 1]
263. ! [X0] : ((v1_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k5_glib_003(X0) = k1_funct_1(X0,k2_glib_003)) [input d7_glib_003]
264. ! [X0] : ((v2_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k6_glib_003(X0) = k1_funct_1(X0,k3_glib_003)) [input d8_glib_003]
265. ! [X0] : ((v3_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k7_glib_003(X0) = k1_funct_1(X0,k4_glib_003)) [input d9_glib_003]
266. ~(! [X0] : ((v3_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => k7_glib_003(X0) = k1_funct_1(X0,k4_glib_003)) & ! [X1] : ((v2_glib_003(X1) & v1_glib_000(X1) & v1_finset_1(X1) & v1_funct_1(X1) & v4_relat_1(X1,k5_numbers) & v1_relat_1(X1)) => k6_glib_003(X1) = k1_funct_1(X1,k3_glib_003)) & ! [X2] : ((v1_glib_003(X2) & v1_glib_000(X2) & v1_finset_1(X2) & v1_funct_1(X2) & v4_relat_1(X2,k5_numbers) & v1_relat_1(X2)) => k5_glib_003(X2) = k1_funct_1(X2,k2_glib_003))) [rectify 2]
346. ? [X0] : (k7_glib_003(X0) != k1_funct_1(X0,k4_glib_003) & (v3_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0))) | ? [X1] : (k6_glib_003(X1) != k1_funct_1(X1,k3_glib_003) & (v2_glib_003(X1) & v1_glib_000(X1) & v1_finset_1(X1) & v1_funct_1(X1) & v4_relat_1(X1,k5_numbers) & v1_relat_1(X1))) | ? [X2] : (k5_glib_003(X2) != k1_funct_1(X2,k2_glib_003) & (v1_glib_003(X2) & v1_glib_000(X2) & v1_finset_1(X2) & v1_funct_1(X2) & v4_relat_1(X2,k5_numbers) & v1_relat_1(X2))) [ennf transformation 266]
347. ? [X0] : (k7_glib_003(X0) != k1_funct_1(X0,k4_glib_003) & v3_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) | ? [X1] : (k6_glib_003(X1) != k1_funct_1(X1,k3_glib_003) & v2_glib_003(X1) & v1_glib_000(X1) & v1_finset_1(X1) & v1_funct_1(X1) & v4_relat_1(X1,k5_numbers) & v1_relat_1(X1)) | ? [X2] : (k5_glib_003(X2) != k1_funct_1(X2,k2_glib_003) & v1_glib_003(X2) & v1_glib_000(X2) & v1_finset_1(X2) & v1_funct_1(X2) & v4_relat_1(X2,k5_numbers) & v1_relat_1(X2)) [flattening 346]
499. ! [X0] : (k6_glib_003(X0) = k1_funct_1(X0,k3_glib_003) | (~v2_glib_003(X0) | ~v1_glib_000(X0) | ~v1_finset_1(X0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k5_numbers) | ~v1_relat_1(X0))) [ennf transformation 264]
500. ! [X0] : (k6_glib_003(X0) = k1_funct_1(X0,k3_glib_003) | ~v2_glib_003(X0) | ~v1_glib_000(X0) | ~v1_finset_1(X0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k5_numbers) | ~v1_relat_1(X0)) [flattening 499]
501. ! [X0] : (k7_glib_003(X0) = k1_funct_1(X0,k4_glib_003) | (~v3_glib_003(X0) | ~v1_glib_000(X0) | ~v1_finset_1(X0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k5_numbers) | ~v1_relat_1(X0))) [ennf transformation 265]
502. ! [X0] : (k7_glib_003(X0) = k1_funct_1(X0,k4_glib_003) | ~v3_glib_003(X0) | ~v1_glib_000(X0) | ~v1_finset_1(X0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k5_numbers) | ~v1_relat_1(X0)) [flattening 501]
503. ! [X0] : (k5_glib_003(X0) = k1_funct_1(X0,k2_glib_003) | (~v1_glib_003(X0) | ~v1_glib_000(X0) | ~v1_finset_1(X0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k5_numbers) | ~v1_relat_1(X0))) [ennf transformation 263]
504. ! [X0] : (k5_glib_003(X0) = k1_funct_1(X0,k2_glib_003) | ~v1_glib_003(X0) | ~v1_glib_000(X0) | ~v1_finset_1(X0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k5_numbers) | ~v1_relat_1(X0)) [flattening 503]
553. ? [X0] : (k7_glib_003(X0) != k1_funct_1(X0,k4_glib_003) & v3_glib_003(X0) & v1_glib_000(X0) & v1_finset_1(X0) & v1_funct_1(X0) & v4_relat_1(X0,k5_numbers) & v1_relat_1(X0)) => (k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) & v3_glib_003(sK0) & v1_glib_000(sK0) & v1_finset_1(sK0) & v1_funct_1(sK0) & v4_relat_1(sK0,k5_numbers) & v1_relat_1(sK0)) [choice axiom]
554. ? [X1] : (k6_glib_003(X1) != k1_funct_1(X1,k3_glib_003) & v2_glib_003(X1) & v1_glib_000(X1) & v1_finset_1(X1) & v1_funct_1(X1) & v4_relat_1(X1,k5_numbers) & v1_relat_1(X1)) => (k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) & v2_glib_003(sK1) & v1_glib_000(sK1) & v1_finset_1(sK1) & v1_funct_1(sK1) & v4_relat_1(sK1,k5_numbers) & v1_relat_1(sK1)) [choice axiom]
555. ? [X2] : (k5_glib_003(X2) != k1_funct_1(X2,k2_glib_003) & v1_glib_003(X2) & v1_glib_000(X2) & v1_finset_1(X2) & v1_funct_1(X2) & v4_relat_1(X2,k5_numbers) & v1_relat_1(X2)) => (k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) & v1_glib_003(sK2) & v1_glib_000(sK2) & v1_finset_1(sK2) & v1_funct_1(sK2) & v4_relat_1(sK2,k5_numbers) & v1_relat_1(sK2)) [choice axiom]
556. (k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) & v3_glib_003(sK0) & v1_glib_000(sK0) & v1_finset_1(sK0) & v1_funct_1(sK0) & v4_relat_1(sK0,k5_numbers) & v1_relat_1(sK0)) | (k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) & v2_glib_003(sK1) & v1_glib_000(sK1) & v1_finset_1(sK1) & v1_funct_1(sK1) & v4_relat_1(sK1,k5_numbers) & v1_relat_1(sK1)) | (k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) & v1_glib_003(sK2) & v1_glib_000(sK2) & v1_finset_1(sK2) & v1_funct_1(sK2) & v4_relat_1(sK2,k5_numbers) & v1_relat_1(sK2)) [skolemisation 347,555,554,553]
676. v1_relat_1(sK0) | v1_relat_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
677. v1_relat_1(sK0) | v1_relat_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
678. v1_relat_1(sK0) | v1_relat_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
679. v1_relat_1(sK0) | v1_relat_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
680. v1_relat_1(sK0) | v1_relat_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
681. v1_relat_1(sK0) | v1_relat_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
682. v1_relat_1(sK0) | v1_relat_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
683. v1_relat_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
684. v1_relat_1(sK0) | v4_relat_1(sK1,k5_numbers) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
685. v1_relat_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
686. v1_relat_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
687. v1_relat_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
688. v1_relat_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
689. v1_relat_1(sK0) | v4_relat_1(sK1,k5_numbers) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
690. v1_relat_1(sK0) | v1_funct_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
691. v1_relat_1(sK0) | v1_funct_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
692. v1_relat_1(sK0) | v1_funct_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
693. v1_relat_1(sK0) | v1_funct_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
694. v1_relat_1(sK0) | v1_funct_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
695. v1_relat_1(sK0) | v1_funct_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
696. v1_relat_1(sK0) | v1_funct_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
697. v1_relat_1(sK0) | v1_finset_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
698. v1_relat_1(sK0) | v1_finset_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
699. v1_relat_1(sK0) | v1_finset_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
700. v1_relat_1(sK0) | v1_finset_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
701. v1_relat_1(sK0) | v1_finset_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
702. v1_relat_1(sK0) | v1_finset_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
703. v1_relat_1(sK0) | v1_finset_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
704. v1_relat_1(sK0) | v1_glib_000(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
705. v1_relat_1(sK0) | v1_glib_000(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
706. v1_relat_1(sK0) | v1_glib_000(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
707. v1_relat_1(sK0) | v1_glib_000(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
708. v1_relat_1(sK0) | v1_glib_000(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
709. v1_relat_1(sK0) | v1_glib_000(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
710. v1_relat_1(sK0) | v1_glib_000(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
711. v1_relat_1(sK0) | v2_glib_003(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
712. v1_relat_1(sK0) | v2_glib_003(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
713. v1_relat_1(sK0) | v2_glib_003(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
714. v1_relat_1(sK0) | v2_glib_003(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
715. v1_relat_1(sK0) | v2_glib_003(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
716. v1_relat_1(sK0) | v2_glib_003(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
717. v1_relat_1(sK0) | v2_glib_003(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
718. v1_relat_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
719. v1_relat_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
720. v1_relat_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
721. v1_relat_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
722. v1_relat_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
723. v1_relat_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
724. v1_relat_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
725. v4_relat_1(sK0,k5_numbers) | v1_relat_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
726. v4_relat_1(sK0,k5_numbers) | v1_relat_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
727. v4_relat_1(sK0,k5_numbers) | v1_relat_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
728. v4_relat_1(sK0,k5_numbers) | v1_relat_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
729. v4_relat_1(sK0,k5_numbers) | v1_relat_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
730. v4_relat_1(sK0,k5_numbers) | v1_relat_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
731. v4_relat_1(sK0,k5_numbers) | v1_relat_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
732. v4_relat_1(sK0,k5_numbers) | v4_relat_1(sK1,k5_numbers) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
733. v4_relat_1(sK0,k5_numbers) | v4_relat_1(sK1,k5_numbers) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
734. v4_relat_1(sK0,k5_numbers) | v4_relat_1(sK1,k5_numbers) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
735. v4_relat_1(sK0,k5_numbers) | v4_relat_1(sK1,k5_numbers) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
736. v4_relat_1(sK0,k5_numbers) | v4_relat_1(sK1,k5_numbers) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
737. v4_relat_1(sK0,k5_numbers) | v4_relat_1(sK1,k5_numbers) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
738. v4_relat_1(sK0,k5_numbers) | v4_relat_1(sK1,k5_numbers) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
739. v4_relat_1(sK0,k5_numbers) | v1_funct_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
740. v4_relat_1(sK0,k5_numbers) | v1_funct_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
741. v4_relat_1(sK0,k5_numbers) | v1_funct_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
742. v4_relat_1(sK0,k5_numbers) | v1_funct_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
743. v4_relat_1(sK0,k5_numbers) | v1_funct_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
744. v4_relat_1(sK0,k5_numbers) | v1_funct_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
745. v4_relat_1(sK0,k5_numbers) | v1_funct_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
746. v4_relat_1(sK0,k5_numbers) | v1_finset_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
747. v4_relat_1(sK0,k5_numbers) | v1_finset_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
748. v4_relat_1(sK0,k5_numbers) | v1_finset_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
749. v4_relat_1(sK0,k5_numbers) | v1_finset_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
750. v4_relat_1(sK0,k5_numbers) | v1_finset_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
751. v4_relat_1(sK0,k5_numbers) | v1_finset_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
752. v4_relat_1(sK0,k5_numbers) | v1_finset_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
753. v4_relat_1(sK0,k5_numbers) | v1_glib_000(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
754. v4_relat_1(sK0,k5_numbers) | v1_glib_000(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
755. v4_relat_1(sK0,k5_numbers) | v1_glib_000(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
756. v4_relat_1(sK0,k5_numbers) | v1_glib_000(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
757. v4_relat_1(sK0,k5_numbers) | v1_glib_000(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
758. v4_relat_1(sK0,k5_numbers) | v1_glib_000(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
759. v4_relat_1(sK0,k5_numbers) | v1_glib_000(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
760. v4_relat_1(sK0,k5_numbers) | v2_glib_003(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
761. v4_relat_1(sK0,k5_numbers) | v2_glib_003(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
762. v4_relat_1(sK0,k5_numbers) | v2_glib_003(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
763. v4_relat_1(sK0,k5_numbers) | v2_glib_003(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
764. v4_relat_1(sK0,k5_numbers) | v2_glib_003(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
765. v4_relat_1(sK0,k5_numbers) | v2_glib_003(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
766. v4_relat_1(sK0,k5_numbers) | v2_glib_003(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
767. v4_relat_1(sK0,k5_numbers) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
768. v4_relat_1(sK0,k5_numbers) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
769. v4_relat_1(sK0,k5_numbers) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
770. v4_relat_1(sK0,k5_numbers) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
771. v4_relat_1(sK0,k5_numbers) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
772. v4_relat_1(sK0,k5_numbers) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
773. v4_relat_1(sK0,k5_numbers) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
774. v1_funct_1(sK0) | v1_relat_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
775. v1_funct_1(sK0) | v1_relat_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
776. v1_funct_1(sK0) | v1_relat_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
777. v1_funct_1(sK0) | v1_relat_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
778. v1_funct_1(sK0) | v1_relat_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
779. v1_funct_1(sK0) | v1_relat_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
780. v1_funct_1(sK0) | v1_relat_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
781. v1_funct_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
782. v1_funct_1(sK0) | v4_relat_1(sK1,k5_numbers) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
783. v1_funct_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
784. v1_funct_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
785. v1_funct_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
786. v1_funct_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
787. v1_funct_1(sK0) | v4_relat_1(sK1,k5_numbers) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
788. v1_funct_1(sK0) | v1_funct_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
789. v1_funct_1(sK0) | v1_funct_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
790. v1_funct_1(sK0) | v1_funct_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
791. v1_funct_1(sK0) | v1_funct_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
792. v1_funct_1(sK0) | v1_funct_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
793. v1_funct_1(sK0) | v1_funct_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
794. v1_funct_1(sK0) | v1_funct_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
795. v1_funct_1(sK0) | v1_finset_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
796. v1_funct_1(sK0) | v1_finset_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
797. v1_funct_1(sK0) | v1_finset_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
798. v1_funct_1(sK0) | v1_finset_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
799. v1_funct_1(sK0) | v1_finset_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
800. v1_funct_1(sK0) | v1_finset_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
801. v1_funct_1(sK0) | v1_finset_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
802. v1_funct_1(sK0) | v1_glib_000(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
803. v1_funct_1(sK0) | v1_glib_000(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
804. v1_funct_1(sK0) | v1_glib_000(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
805. v1_funct_1(sK0) | v1_glib_000(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
806. v1_funct_1(sK0) | v1_glib_000(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
807. v1_funct_1(sK0) | v1_glib_000(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
808. v1_funct_1(sK0) | v1_glib_000(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
809. v1_funct_1(sK0) | v2_glib_003(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
810. v1_funct_1(sK0) | v2_glib_003(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
811. v1_funct_1(sK0) | v2_glib_003(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
812. v1_funct_1(sK0) | v2_glib_003(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
813. v1_funct_1(sK0) | v2_glib_003(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
814. v1_funct_1(sK0) | v2_glib_003(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
815. v1_funct_1(sK0) | v2_glib_003(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
816. v1_funct_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
817. v1_funct_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
818. v1_funct_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
819. v1_funct_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
820. v1_funct_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
821. v1_funct_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
822. v1_funct_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
823. v1_finset_1(sK0) | v1_relat_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
824. v1_finset_1(sK0) | v1_relat_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
825. v1_finset_1(sK0) | v1_relat_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
826. v1_finset_1(sK0) | v1_relat_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
827. v1_finset_1(sK0) | v1_relat_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
828. v1_finset_1(sK0) | v1_relat_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
829. v1_finset_1(sK0) | v1_relat_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
830. v1_finset_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
831. v1_finset_1(sK0) | v4_relat_1(sK1,k5_numbers) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
832. v1_finset_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
833. v1_finset_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
834. v1_finset_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
835. v1_finset_1(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
836. v1_finset_1(sK0) | v4_relat_1(sK1,k5_numbers) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
837. v1_finset_1(sK0) | v1_funct_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
838. v1_finset_1(sK0) | v1_funct_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
839. v1_finset_1(sK0) | v1_funct_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
840. v1_finset_1(sK0) | v1_funct_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
841. v1_finset_1(sK0) | v1_funct_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
842. v1_finset_1(sK0) | v1_funct_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
843. v1_finset_1(sK0) | v1_funct_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
844. v1_finset_1(sK0) | v1_finset_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
845. v1_finset_1(sK0) | v1_finset_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
846. v1_finset_1(sK0) | v1_finset_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
847. v1_finset_1(sK0) | v1_finset_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
848. v1_finset_1(sK0) | v1_finset_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
849. v1_finset_1(sK0) | v1_finset_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
850. v1_finset_1(sK0) | v1_finset_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
851. v1_finset_1(sK0) | v1_glib_000(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
852. v1_finset_1(sK0) | v1_glib_000(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
853. v1_finset_1(sK0) | v1_glib_000(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
854. v1_finset_1(sK0) | v1_glib_000(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
855. v1_finset_1(sK0) | v1_glib_000(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
856. v1_finset_1(sK0) | v1_glib_000(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
857. v1_finset_1(sK0) | v1_glib_000(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
858. v1_finset_1(sK0) | v2_glib_003(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
859. v1_finset_1(sK0) | v2_glib_003(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
860. v1_finset_1(sK0) | v2_glib_003(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
861. v1_finset_1(sK0) | v2_glib_003(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
862. v1_finset_1(sK0) | v2_glib_003(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
863. v1_finset_1(sK0) | v2_glib_003(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
864. v1_finset_1(sK0) | v2_glib_003(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
865. v1_finset_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
866. v1_finset_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
867. v1_finset_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
868. v1_finset_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
869. v1_finset_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
870. v1_finset_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
871. v1_finset_1(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
872. v1_glib_000(sK0) | v1_relat_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
873. v1_glib_000(sK0) | v1_relat_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
874. v1_glib_000(sK0) | v1_relat_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
875. v1_glib_000(sK0) | v1_relat_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
876. v1_glib_000(sK0) | v1_relat_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
877. v1_glib_000(sK0) | v1_relat_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
878. v1_glib_000(sK0) | v1_relat_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
879. v1_glib_000(sK0) | v4_relat_1(sK1,k5_numbers) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
880. v1_glib_000(sK0) | v4_relat_1(sK1,k5_numbers) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
881. v1_glib_000(sK0) | v4_relat_1(sK1,k5_numbers) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
882. v1_glib_000(sK0) | v4_relat_1(sK1,k5_numbers) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
883. v1_glib_000(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
884. v1_glib_000(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
885. v1_glib_000(sK0) | v4_relat_1(sK1,k5_numbers) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
886. v1_glib_000(sK0) | v1_funct_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
887. v1_glib_000(sK0) | v1_funct_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
888. v1_glib_000(sK0) | v1_funct_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
889. v1_glib_000(sK0) | v1_funct_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
890. v1_glib_000(sK0) | v1_funct_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
891. v1_glib_000(sK0) | v1_funct_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
892. v1_glib_000(sK0) | v1_funct_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
893. v1_glib_000(sK0) | v1_finset_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
894. v1_glib_000(sK0) | v1_finset_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
895. v1_glib_000(sK0) | v1_finset_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
896. v1_glib_000(sK0) | v1_finset_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
897. v1_glib_000(sK0) | v1_finset_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
898. v1_glib_000(sK0) | v1_finset_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
899. v1_glib_000(sK0) | v1_finset_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
900. v1_glib_000(sK0) | v1_glib_000(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
901. v1_glib_000(sK0) | v1_glib_000(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
902. v1_glib_000(sK0) | v1_glib_000(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
903. v1_glib_000(sK0) | v1_glib_000(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
904. v1_glib_000(sK0) | v1_glib_000(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
905. v1_glib_000(sK0) | v1_glib_000(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
906. v1_glib_000(sK0) | v1_glib_000(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
907. v1_glib_000(sK0) | v2_glib_003(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
908. v1_glib_000(sK0) | v2_glib_003(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
909. v1_glib_000(sK0) | v2_glib_003(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
910. v1_glib_000(sK0) | v2_glib_003(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
911. v1_glib_000(sK0) | v2_glib_003(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
912. v1_glib_000(sK0) | v2_glib_003(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
913. v1_glib_000(sK0) | v2_glib_003(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
914. v1_glib_000(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
915. v1_glib_000(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
916. v1_glib_000(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
917. v1_glib_000(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
918. v1_glib_000(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
919. v1_glib_000(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
920. v1_glib_000(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
921. v3_glib_003(sK0) | v1_relat_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
922. v3_glib_003(sK0) | v1_relat_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
923. v3_glib_003(sK0) | v1_relat_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
924. v3_glib_003(sK0) | v1_relat_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
925. v3_glib_003(sK0) | v1_relat_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
926. v3_glib_003(sK0) | v1_relat_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
927. v3_glib_003(sK0) | v1_relat_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
928. v3_glib_003(sK0) | v4_relat_1(sK1,k5_numbers) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
929. v3_glib_003(sK0) | v4_relat_1(sK1,k5_numbers) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
930. v3_glib_003(sK0) | v4_relat_1(sK1,k5_numbers) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
931. v3_glib_003(sK0) | v4_relat_1(sK1,k5_numbers) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
932. v3_glib_003(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
933. v3_glib_003(sK0) | v4_relat_1(sK1,k5_numbers) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
934. v3_glib_003(sK0) | v4_relat_1(sK1,k5_numbers) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
935. v3_glib_003(sK0) | v1_funct_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
936. v3_glib_003(sK0) | v1_funct_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
937. v3_glib_003(sK0) | v1_funct_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
938. v3_glib_003(sK0) | v1_funct_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
939. v3_glib_003(sK0) | v1_funct_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
940. v3_glib_003(sK0) | v1_funct_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
941. v3_glib_003(sK0) | v1_funct_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
942. v3_glib_003(sK0) | v1_finset_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
943. v3_glib_003(sK0) | v1_finset_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
944. v3_glib_003(sK0) | v1_finset_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
945. v3_glib_003(sK0) | v1_finset_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
946. v3_glib_003(sK0) | v1_finset_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
947. v3_glib_003(sK0) | v1_finset_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
948. v3_glib_003(sK0) | v1_finset_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
949. v3_glib_003(sK0) | v1_glib_000(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
950. v3_glib_003(sK0) | v1_glib_000(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
951. v3_glib_003(sK0) | v1_glib_000(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
952. v3_glib_003(sK0) | v1_glib_000(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
953. v3_glib_003(sK0) | v1_glib_000(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
954. v3_glib_003(sK0) | v1_glib_000(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
955. v3_glib_003(sK0) | v1_glib_000(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
956. v3_glib_003(sK0) | v2_glib_003(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
957. v3_glib_003(sK0) | v2_glib_003(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
958. v3_glib_003(sK0) | v2_glib_003(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
959. v3_glib_003(sK0) | v2_glib_003(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
960. v3_glib_003(sK0) | v2_glib_003(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
961. v3_glib_003(sK0) | v2_glib_003(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
962. v3_glib_003(sK0) | v2_glib_003(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
963. v3_glib_003(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
964. v3_glib_003(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
965. v3_glib_003(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
966. v3_glib_003(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
967. v3_glib_003(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
968. v3_glib_003(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
969. v3_glib_003(sK0) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
970. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_relat_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
971. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_relat_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
972. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_relat_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
973. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_relat_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
974. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_relat_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
975. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_relat_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
976. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_relat_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
977. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v4_relat_1(sK1,k5_numbers) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
978. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v4_relat_1(sK1,k5_numbers) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
979. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v4_relat_1(sK1,k5_numbers) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
980. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v4_relat_1(sK1,k5_numbers) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
981. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v4_relat_1(sK1,k5_numbers) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
982. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v4_relat_1(sK1,k5_numbers) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
983. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v4_relat_1(sK1,k5_numbers) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
984. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_funct_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
985. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_funct_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
986. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_funct_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
987. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_funct_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
988. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_funct_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
989. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_funct_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
990. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_funct_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
991. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_finset_1(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
992. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_finset_1(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
993. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_finset_1(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
994. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_finset_1(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
995. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_finset_1(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
996. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_finset_1(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
997. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_finset_1(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
998. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_glib_000(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
999. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_glib_000(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1000. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_glib_000(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1001. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_glib_000(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1002. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_glib_000(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1003. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_glib_000(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1004. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v1_glib_000(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1005. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v2_glib_003(sK1) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1006. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v2_glib_003(sK1) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1007. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v2_glib_003(sK1) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1008. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v2_glib_003(sK1) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1009. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v2_glib_003(sK1) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1010. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v2_glib_003(sK1) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1011. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | v2_glib_003(sK1) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1012. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_relat_1(sK2) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1013. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v4_relat_1(sK2,k5_numbers) [cnf transformation 556] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1014. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_funct_1(sK2) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1015. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_finset_1(sK2) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1016. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_000(sK2) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1017. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | v1_glib_003(sK2) [cnf transformation 556] {a:0,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1018. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) | k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) | k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) [cnf transformation 556] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
1247. k6_glib_003(X0) = k1_funct_1(X0,k3_glib_003) | ~v2_glib_003(X0) | ~v1_glib_000(X0) | ~v1_finset_1(X0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k5_numbers) | ~v1_relat_1(X0) [cnf transformation 500] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1248. k7_glib_003(X0) = k1_funct_1(X0,k4_glib_003) | ~v3_glib_003(X0) | ~v1_glib_000(X0) | ~v1_finset_1(X0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k5_numbers) | ~v1_relat_1(X0) [cnf transformation 502] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1249. k5_glib_003(X0) = k1_funct_1(X0,k2_glib_003) | ~v1_glib_003(X0) | ~v1_glib_000(X0) | ~v1_finset_1(X0) | ~v1_funct_1(X0) | ~v4_relat_1(X0,k5_numbers) | ~v1_relat_1(X0) [cnf transformation 504] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1455. 1 <=> k5_glib_003(sK2) = k1_funct_1(sK2,k2_glib_003) [avatar definition]
1457. k5_glib_003(sK2) != k1_funct_1(sK2,k2_glib_003) <- (~1) [avatar component clause 1455] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1459. 2 <=> k6_glib_003(sK1) = k1_funct_1(sK1,k3_glib_003) [avatar definition]
1461. k6_glib_003(sK1) != k1_funct_1(sK1,k3_glib_003) <- (~2) [avatar component clause 1459] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1463. 3 <=> k7_glib_003(sK0) = k1_funct_1(sK0,k4_glib_003) [avatar definition]
1465. k7_glib_003(sK0) != k1_funct_1(sK0,k4_glib_003) <- (~3) [avatar component clause 1463] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1466. ~1 | ~2 | ~3 [avatar split clause 1018,1463,1459,1455]
1468. 4 <=> v1_glib_003(sK2) [avatar definition]
1470. v1_glib_003(sK2) <- (4) [avatar component clause 1468] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1471. 4 | ~2 | ~3 [avatar split clause 1017,1463,1459,1468]
1473. 5 <=> v1_glib_000(sK2) [avatar definition]
1475. v1_glib_000(sK2) <- (5) [avatar component clause 1473] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1476. 5 | ~2 | ~3 [avatar split clause 1016,1463,1459,1473]
1478. 6 <=> v1_finset_1(sK2) [avatar definition]
1480. v1_finset_1(sK2) <- (6) [avatar component clause 1478] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1481. 6 | ~2 | ~3 [avatar split clause 1015,1463,1459,1478]
1483. 7 <=> v1_funct_1(sK2) [avatar definition]
1485. v1_funct_1(sK2) <- (7) [avatar component clause 1483] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1486. 7 | ~2 | ~3 [avatar split clause 1014,1463,1459,1483]
1488. 8 <=> v4_relat_1(sK2,k5_numbers) [avatar definition]
1490. v4_relat_1(sK2,k5_numbers) <- (8) [avatar component clause 1488] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1491. 8 | ~2 | ~3 [avatar split clause 1013,1463,1459,1488]
1493. 9 <=> v1_relat_1(sK2) [avatar definition]
1495. v1_relat_1(sK2) <- (9) [avatar component clause 1493] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1496. 9 | ~2 | ~3 [avatar split clause 1012,1463,1459,1493]
1498. 10 <=> v2_glib_003(sK1) [avatar definition]
1500. v2_glib_003(sK1) <- (10) [avatar component clause 1498] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1501. ~1 | 10 | ~3 [avatar split clause 1011,1463,1498,1455]
1502. 4 | 10 | ~3 [avatar split clause 1010,1463,1498,1468]
1503. 5 | 10 | ~3 [avatar split clause 1009,1463,1498,1473]
1504. 6 | 10 | ~3 [avatar split clause 1008,1463,1498,1478]
1505. 7 | 10 | ~3 [avatar split clause 1007,1463,1498,1483]
1506. 8 | 10 | ~3 [avatar split clause 1006,1463,1498,1488]
1507. 9 | 10 | ~3 [avatar split clause 1005,1463,1498,1493]
1509. 11 <=> v1_glib_000(sK1) [avatar definition]
1511. v1_glib_000(sK1) <- (11) [avatar component clause 1509] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1512. ~1 | 11 | ~3 [avatar split clause 1004,1463,1509,1455]
1513. 4 | 11 | ~3 [avatar split clause 1003,1463,1509,1468]
1514. 5 | 11 | ~3 [avatar split clause 1002,1463,1509,1473]
1515. 6 | 11 | ~3 [avatar split clause 1001,1463,1509,1478]
1516. 7 | 11 | ~3 [avatar split clause 1000,1463,1509,1483]
1517. 8 | 11 | ~3 [avatar split clause 999,1463,1509,1488]
1518. 9 | 11 | ~3 [avatar split clause 998,1463,1509,1493]
1520. 12 <=> v1_finset_1(sK1) [avatar definition]
1522. v1_finset_1(sK1) <- (12) [avatar component clause 1520] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1523. ~1 | 12 | ~3 [avatar split clause 997,1463,1520,1455]
1524. 4 | 12 | ~3 [avatar split clause 996,1463,1520,1468]
1525. 5 | 12 | ~3 [avatar split clause 995,1463,1520,1473]
1526. 6 | 12 | ~3 [avatar split clause 994,1463,1520,1478]
1527. 7 | 12 | ~3 [avatar split clause 993,1463,1520,1483]
1528. 8 | 12 | ~3 [avatar split clause 992,1463,1520,1488]
1529. 9 | 12 | ~3 [avatar split clause 991,1463,1520,1493]
1531. 13 <=> v1_funct_1(sK1) [avatar definition]
1533. v1_funct_1(sK1) <- (13) [avatar component clause 1531] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1534. ~1 | 13 | ~3 [avatar split clause 990,1463,1531,1455]
1535. 4 | 13 | ~3 [avatar split clause 989,1463,1531,1468]
1536. 5 | 13 | ~3 [avatar split clause 988,1463,1531,1473]
1537. 6 | 13 | ~3 [avatar split clause 987,1463,1531,1478]
1538. 7 | 13 | ~3 [avatar split clause 986,1463,1531,1483]
1539. 8 | 13 | ~3 [avatar split clause 985,1463,1531,1488]
1540. 9 | 13 | ~3 [avatar split clause 984,1463,1531,1493]
1542. 14 <=> v4_relat_1(sK1,k5_numbers) [avatar definition]
1544. v4_relat_1(sK1,k5_numbers) <- (14) [avatar component clause 1542] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1545. ~1 | 14 | ~3 [avatar split clause 983,1463,1542,1455]
1546. 4 | 14 | ~3 [avatar split clause 982,1463,1542,1468]
1547. 5 | 14 | ~3 [avatar split clause 981,1463,1542,1473]
1548. 6 | 14 | ~3 [avatar split clause 980,1463,1542,1478]
1549. 7 | 14 | ~3 [avatar split clause 979,1463,1542,1483]
1550. 8 | 14 | ~3 [avatar split clause 978,1463,1542,1488]
1551. 9 | 14 | ~3 [avatar split clause 977,1463,1542,1493]
1553. 15 <=> v1_relat_1(sK1) [avatar definition]
1556. ~1 | 15 | ~3 [avatar split clause 976,1463,1553,1455]
1557. 4 | 15 | ~3 [avatar split clause 975,1463,1553,1468]
1558. 5 | 15 | ~3 [avatar split clause 974,1463,1553,1473]
1559. 6 | 15 | ~3 [avatar split clause 973,1463,1553,1478]
1560. 7 | 15 | ~3 [avatar split clause 972,1463,1553,1483]
1561. 8 | 15 | ~3 [avatar split clause 971,1463,1553,1488]
1562. 9 | 15 | ~3 [avatar split clause 970,1463,1553,1493]
1564. 16 <=> v3_glib_003(sK0) [avatar definition]
1566. v3_glib_003(sK0) <- (16) [avatar component clause 1564] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1567. ~1 | ~2 | 16 [avatar split clause 969,1564,1459,1455]
1568. 4 | ~2 | 16 [avatar split clause 968,1564,1459,1468]
1569. 5 | ~2 | 16 [avatar split clause 967,1564,1459,1473]
1570. 6 | ~2 | 16 [avatar split clause 966,1564,1459,1478]
1571. 7 | ~2 | 16 [avatar split clause 965,1564,1459,1483]
1572. 8 | ~2 | 16 [avatar split clause 964,1564,1459,1488]
1573. 9 | ~2 | 16 [avatar split clause 963,1564,1459,1493]
1574. ~1 | 10 | 16 [avatar split clause 962,1564,1498,1455]
1575. 4 | 10 | 16 [avatar split clause 961,1564,1498,1468]
1576. 5 | 10 | 16 [avatar split clause 960,1564,1498,1473]
1577. 6 | 10 | 16 [avatar split clause 959,1564,1498,1478]
1578. 7 | 10 | 16 [avatar split clause 958,1564,1498,1483]
1579. 8 | 10 | 16 [avatar split clause 957,1564,1498,1488]
1580. 9 | 10 | 16 [avatar split clause 956,1564,1498,1493]
1581. ~1 | 11 | 16 [avatar split clause 955,1564,1509,1455]
1582. 4 | 11 | 16 [avatar split clause 954,1564,1509,1468]
1583. 5 | 11 | 16 [avatar split clause 953,1564,1509,1473]
1584. 6 | 11 | 16 [avatar split clause 952,1564,1509,1478]
1585. 7 | 11 | 16 [avatar split clause 951,1564,1509,1483]
1586. 8 | 11 | 16 [avatar split clause 950,1564,1509,1488]
1587. 9 | 11 | 16 [avatar split clause 949,1564,1509,1493]
1588. ~1 | 12 | 16 [avatar split clause 948,1564,1520,1455]
1589. 4 | 12 | 16 [avatar split clause 947,1564,1520,1468]
1590. 5 | 12 | 16 [avatar split clause 946,1564,1520,1473]
1591. 6 | 12 | 16 [avatar split clause 945,1564,1520,1478]
1592. 7 | 12 | 16 [avatar split clause 944,1564,1520,1483]
1593. 8 | 12 | 16 [avatar split clause 943,1564,1520,1488]
1594. 9 | 12 | 16 [avatar split clause 942,1564,1520,1493]
1595. ~1 | 13 | 16 [avatar split clause 941,1564,1531,1455]
1596. 4 | 13 | 16 [avatar split clause 940,1564,1531,1468]
1597. 5 | 13 | 16 [avatar split clause 939,1564,1531,1473]
1598. 6 | 13 | 16 [avatar split clause 938,1564,1531,1478]
1599. 7 | 13 | 16 [avatar split clause 937,1564,1531,1483]
1600. 8 | 13 | 16 [avatar split clause 936,1564,1531,1488]
1601. 9 | 13 | 16 [avatar split clause 935,1564,1531,1493]
1602. ~1 | 14 | 16 [avatar split clause 934,1564,1542,1455]
1603. 4 | 14 | 16 [avatar split clause 933,1564,1542,1468]
1604. 5 | 14 | 16 [avatar split clause 932,1564,1542,1473]
1605. 6 | 14 | 16 [avatar split clause 931,1564,1542,1478]
1606. 7 | 14 | 16 [avatar split clause 930,1564,1542,1483]
1607. 8 | 14 | 16 [avatar split clause 929,1564,1542,1488]
1608. 9 | 14 | 16 [avatar split clause 928,1564,1542,1493]
1609. ~1 | 15 | 16 [avatar split clause 927,1564,1553,1455]
1610. 4 | 15 | 16 [avatar split clause 926,1564,1553,1468]
1611. 5 | 15 | 16 [avatar split clause 925,1564,1553,1473]
1612. 6 | 15 | 16 [avatar split clause 924,1564,1553,1478]
1613. 7 | 15 | 16 [avatar split clause 923,1564,1553,1483]
1614. 8 | 15 | 16 [avatar split clause 922,1564,1553,1488]
1615. 9 | 15 | 16 [avatar split clause 921,1564,1553,1493]
1617. 17 <=> v1_glib_000(sK0) [avatar definition]
1619. v1_glib_000(sK0) <- (17) [avatar component clause 1617] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1620. ~1 | ~2 | 17 [avatar split clause 920,1617,1459,1455]
1621. 4 | ~2 | 17 [avatar split clause 919,1617,1459,1468]
1622. 5 | ~2 | 17 [avatar split clause 918,1617,1459,1473]
1623. 6 | ~2 | 17 [avatar split clause 917,1617,1459,1478]
1624. 7 | ~2 | 17 [avatar split clause 916,1617,1459,1483]
1625. 8 | ~2 | 17 [avatar split clause 915,1617,1459,1488]
1626. 9 | ~2 | 17 [avatar split clause 914,1617,1459,1493]
1627. ~1 | 10 | 17 [avatar split clause 913,1617,1498,1455]
1628. 4 | 10 | 17 [avatar split clause 912,1617,1498,1468]
1629. 5 | 10 | 17 [avatar split clause 911,1617,1498,1473]
1630. 6 | 10 | 17 [avatar split clause 910,1617,1498,1478]
1631. 7 | 10 | 17 [avatar split clause 909,1617,1498,1483]
1632. 8 | 10 | 17 [avatar split clause 908,1617,1498,1488]
1633. 9 | 10 | 17 [avatar split clause 907,1617,1498,1493]
1634. ~1 | 11 | 17 [avatar split clause 906,1617,1509,1455]
1635. 4 | 11 | 17 [avatar split clause 905,1617,1509,1468]
1636. 5 | 11 | 17 [avatar split clause 904,1617,1509,1473]
1637. 6 | 11 | 17 [avatar split clause 903,1617,1509,1478]
1638. 7 | 11 | 17 [avatar split clause 902,1617,1509,1483]
1639. 8 | 11 | 17 [avatar split clause 901,1617,1509,1488]
1640. 9 | 11 | 17 [avatar split clause 900,1617,1509,1493]
1641. ~1 | 12 | 17 [avatar split clause 899,1617,1520,1455]
1642. 4 | 12 | 17 [avatar split clause 898,1617,1520,1468]
1643. 5 | 12 | 17 [avatar split clause 897,1617,1520,1473]
1644. 6 | 12 | 17 [avatar split clause 896,1617,1520,1478]
1645. 7 | 12 | 17 [avatar split clause 895,1617,1520,1483]
1646. 8 | 12 | 17 [avatar split clause 894,1617,1520,1488]
1647. 9 | 12 | 17 [avatar split clause 893,1617,1520,1493]
1648. ~1 | 13 | 17 [avatar split clause 892,1617,1531,1455]
1649. 4 | 13 | 17 [avatar split clause 891,1617,1531,1468]
1650. 5 | 13 | 17 [avatar split clause 890,1617,1531,1473]
1651. 6 | 13 | 17 [avatar split clause 889,1617,1531,1478]
1652. 7 | 13 | 17 [avatar split clause 888,1617,1531,1483]
1653. 8 | 13 | 17 [avatar split clause 887,1617,1531,1488]
1654. 9 | 13 | 17 [avatar split clause 886,1617,1531,1493]
1655. ~1 | 14 | 17 [avatar split clause 885,1617,1542,1455]
1656. 4 | 14 | 17 [avatar split clause 884,1617,1542,1468]
1657. 5 | 14 | 17 [avatar split clause 883,1617,1542,1473]
1658. 6 | 14 | 17 [avatar split clause 882,1617,1542,1478]
1659. 7 | 14 | 17 [avatar split clause 881,1617,1542,1483]
1660. 8 | 14 | 17 [avatar split clause 880,1617,1542,1488]
1661. 9 | 14 | 17 [avatar split clause 879,1617,1542,1493]
1662. ~1 | 15 | 17 [avatar split clause 878,1617,1553,1455]
1663. 4 | 15 | 17 [avatar split clause 877,1617,1553,1468]
1664. 5 | 15 | 17 [avatar split clause 876,1617,1553,1473]
1665. 6 | 15 | 17 [avatar split clause 875,1617,1553,1478]
1666. 7 | 15 | 17 [avatar split clause 874,1617,1553,1483]
1667. 8 | 15 | 17 [avatar split clause 873,1617,1553,1488]
1668. 9 | 15 | 17 [avatar split clause 872,1617,1553,1493]
1670. 18 <=> v1_finset_1(sK0) [avatar definition]
1672. v1_finset_1(sK0) <- (18) [avatar component clause 1670] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1673. ~1 | ~2 | 18 [avatar split clause 871,1670,1459,1455]
1674. 4 | ~2 | 18 [avatar split clause 870,1670,1459,1468]
1675. 5 | ~2 | 18 [avatar split clause 869,1670,1459,1473]
1676. 6 | ~2 | 18 [avatar split clause 868,1670,1459,1478]
1677. 7 | ~2 | 18 [avatar split clause 867,1670,1459,1483]
1678. 8 | ~2 | 18 [avatar split clause 866,1670,1459,1488]
1679. 9 | ~2 | 18 [avatar split clause 865,1670,1459,1493]
1680. ~1 | 10 | 18 [avatar split clause 864,1670,1498,1455]
1681. 4 | 10 | 18 [avatar split clause 863,1670,1498,1468]
1682. 5 | 10 | 18 [avatar split clause 862,1670,1498,1473]
1683. 6 | 10 | 18 [avatar split clause 861,1670,1498,1478]
1684. 7 | 10 | 18 [avatar split clause 860,1670,1498,1483]
1685. 8 | 10 | 18 [avatar split clause 859,1670,1498,1488]
1686. 9 | 10 | 18 [avatar split clause 858,1670,1498,1493]
1687. ~1 | 11 | 18 [avatar split clause 857,1670,1509,1455]
1688. 4 | 11 | 18 [avatar split clause 856,1670,1509,1468]
1689. 5 | 11 | 18 [avatar split clause 855,1670,1509,1473]
1690. 6 | 11 | 18 [avatar split clause 854,1670,1509,1478]
1691. 7 | 11 | 18 [avatar split clause 853,1670,1509,1483]
1692. 8 | 11 | 18 [avatar split clause 852,1670,1509,1488]
1693. 9 | 11 | 18 [avatar split clause 851,1670,1509,1493]
1694. ~1 | 12 | 18 [avatar split clause 850,1670,1520,1455]
1695. 4 | 12 | 18 [avatar split clause 849,1670,1520,1468]
1696. 5 | 12 | 18 [avatar split clause 848,1670,1520,1473]
1697. 6 | 12 | 18 [avatar split clause 847,1670,1520,1478]
1698. 7 | 12 | 18 [avatar split clause 846,1670,1520,1483]
1699. 8 | 12 | 18 [avatar split clause 845,1670,1520,1488]
1700. 9 | 12 | 18 [avatar split clause 844,1670,1520,1493]
1701. ~1 | 13 | 18 [avatar split clause 843,1670,1531,1455]
1702. 4 | 13 | 18 [avatar split clause 842,1670,1531,1468]
1703. 5 | 13 | 18 [avatar split clause 841,1670,1531,1473]
1704. 6 | 13 | 18 [avatar split clause 840,1670,1531,1478]
1705. 7 | 13 | 18 [avatar split clause 839,1670,1531,1483]
1706. 8 | 13 | 18 [avatar split clause 838,1670,1531,1488]
1707. 9 | 13 | 18 [avatar split clause 837,1670,1531,1493]
1708. ~1 | 14 | 18 [avatar split clause 836,1670,1542,1455]
1709. 4 | 14 | 18 [avatar split clause 835,1670,1542,1468]
1710. 5 | 14 | 18 [avatar split clause 834,1670,1542,1473]
1711. 6 | 14 | 18 [avatar split clause 833,1670,1542,1478]
1712. 7 | 14 | 18 [avatar split clause 832,1670,1542,1483]
1713. 8 | 14 | 18 [avatar split clause 831,1670,1542,1488]
1714. 9 | 14 | 18 [avatar split clause 830,1670,1542,1493]
1715. ~1 | 15 | 18 [avatar split clause 829,1670,1553,1455]
1716. 4 | 15 | 18 [avatar split clause 828,1670,1553,1468]
1717. 5 | 15 | 18 [avatar split clause 827,1670,1553,1473]
1718. 6 | 15 | 18 [avatar split clause 826,1670,1553,1478]
1719. 7 | 15 | 18 [avatar split clause 825,1670,1553,1483]
1720. 8 | 15 | 18 [avatar split clause 824,1670,1553,1488]
1721. 9 | 15 | 18 [avatar split clause 823,1670,1553,1493]
1723. 19 <=> v1_funct_1(sK0) [avatar definition]
1725. v1_funct_1(sK0) <- (19) [avatar component clause 1723] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1726. ~1 | ~2 | 19 [avatar split clause 822,1723,1459,1455]
1727. 4 | ~2 | 19 [avatar split clause 821,1723,1459,1468]
1728. 5 | ~2 | 19 [avatar split clause 820,1723,1459,1473]
1729. 6 | ~2 | 19 [avatar split clause 819,1723,1459,1478]
1730. 7 | ~2 | 19 [avatar split clause 818,1723,1459,1483]
1731. 8 | ~2 | 19 [avatar split clause 817,1723,1459,1488]
1732. 9 | ~2 | 19 [avatar split clause 816,1723,1459,1493]
1733. ~1 | 10 | 19 [avatar split clause 815,1723,1498,1455]
1734. 4 | 10 | 19 [avatar split clause 814,1723,1498,1468]
1735. 5 | 10 | 19 [avatar split clause 813,1723,1498,1473]
1736. 6 | 10 | 19 [avatar split clause 812,1723,1498,1478]
1737. 7 | 10 | 19 [avatar split clause 811,1723,1498,1483]
1738. 8 | 10 | 19 [avatar split clause 810,1723,1498,1488]
1739. 9 | 10 | 19 [avatar split clause 809,1723,1498,1493]
1740. ~1 | 11 | 19 [avatar split clause 808,1723,1509,1455]
1741. 4 | 11 | 19 [avatar split clause 807,1723,1509,1468]
1742. 5 | 11 | 19 [avatar split clause 806,1723,1509,1473]
1743. 6 | 11 | 19 [avatar split clause 805,1723,1509,1478]
1744. 7 | 11 | 19 [avatar split clause 804,1723,1509,1483]
1745. 8 | 11 | 19 [avatar split clause 803,1723,1509,1488]
1746. 9 | 11 | 19 [avatar split clause 802,1723,1509,1493]
1747. ~1 | 12 | 19 [avatar split clause 801,1723,1520,1455]
1748. 4 | 12 | 19 [avatar split clause 800,1723,1520,1468]
1749. 5 | 12 | 19 [avatar split clause 799,1723,1520,1473]
1750. 6 | 12 | 19 [avatar split clause 798,1723,1520,1478]
1751. 7 | 12 | 19 [avatar split clause 797,1723,1520,1483]
1752. 8 | 12 | 19 [avatar split clause 796,1723,1520,1488]
1753. 9 | 12 | 19 [avatar split clause 795,1723,1520,1493]
1754. ~1 | 13 | 19 [avatar split clause 794,1723,1531,1455]
1755. 4 | 13 | 19 [avatar split clause 793,1723,1531,1468]
1756. 5 | 13 | 19 [avatar split clause 792,1723,1531,1473]
1757. 6 | 13 | 19 [avatar split clause 791,1723,1531,1478]
1758. 7 | 13 | 19 [avatar split clause 790,1723,1531,1483]
1759. 8 | 13 | 19 [avatar split clause 789,1723,1531,1488]
1760. 9 | 13 | 19 [avatar split clause 788,1723,1531,1493]
1761. ~1 | 14 | 19 [avatar split clause 787,1723,1542,1455]
1762. 4 | 14 | 19 [avatar split clause 786,1723,1542,1468]
1763. 5 | 14 | 19 [avatar split clause 785,1723,1542,1473]
1764. 6 | 14 | 19 [avatar split clause 784,1723,1542,1478]
1765. 7 | 14 | 19 [avatar split clause 783,1723,1542,1483]
1766. 8 | 14 | 19 [avatar split clause 782,1723,1542,1488]
1767. 9 | 14 | 19 [avatar split clause 781,1723,1542,1493]
1768. ~1 | 15 | 19 [avatar split clause 780,1723,1553,1455]
1769. 4 | 15 | 19 [avatar split clause 779,1723,1553,1468]
1770. 5 | 15 | 19 [avatar split clause 778,1723,1553,1473]
1771. 6 | 15 | 19 [avatar split clause 777,1723,1553,1478]
1772. 7 | 15 | 19 [avatar split clause 776,1723,1553,1483]
1773. 8 | 15 | 19 [avatar split clause 775,1723,1553,1488]
1774. 9 | 15 | 19 [avatar split clause 774,1723,1553,1493]
1776. 20 <=> v4_relat_1(sK0,k5_numbers) [avatar definition]
1778. v4_relat_1(sK0,k5_numbers) <- (20) [avatar component clause 1776] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1779. ~1 | ~2 | 20 [avatar split clause 773,1776,1459,1455]
1780. 4 | ~2 | 20 [avatar split clause 772,1776,1459,1468]
1781. 5 | ~2 | 20 [avatar split clause 771,1776,1459,1473]
1782. 6 | ~2 | 20 [avatar split clause 770,1776,1459,1478]
1783. 7 | ~2 | 20 [avatar split clause 769,1776,1459,1483]
1784. 8 | ~2 | 20 [avatar split clause 768,1776,1459,1488]
1785. 9 | ~2 | 20 [avatar split clause 767,1776,1459,1493]
1786. ~1 | 10 | 20 [avatar split clause 766,1776,1498,1455]
1787. 4 | 10 | 20 [avatar split clause 765,1776,1498,1468]
1788. 5 | 10 | 20 [avatar split clause 764,1776,1498,1473]
1789. 6 | 10 | 20 [avatar split clause 763,1776,1498,1478]
1790. 7 | 10 | 20 [avatar split clause 762,1776,1498,1483]
1791. 8 | 10 | 20 [avatar split clause 761,1776,1498,1488]
1792. 9 | 10 | 20 [avatar split clause 760,1776,1498,1493]
1793. ~1 | 11 | 20 [avatar split clause 759,1776,1509,1455]
1794. 4 | 11 | 20 [avatar split clause 758,1776,1509,1468]
1795. 5 | 11 | 20 [avatar split clause 757,1776,1509,1473]
1796. 6 | 11 | 20 [avatar split clause 756,1776,1509,1478]
1797. 7 | 11 | 20 [avatar split clause 755,1776,1509,1483]
1798. 8 | 11 | 20 [avatar split clause 754,1776,1509,1488]
1799. 9 | 11 | 20 [avatar split clause 753,1776,1509,1493]
1800. ~1 | 12 | 20 [avatar split clause 752,1776,1520,1455]
1801. 4 | 12 | 20 [avatar split clause 751,1776,1520,1468]
1802. 5 | 12 | 20 [avatar split clause 750,1776,1520,1473]
1803. 6 | 12 | 20 [avatar split clause 749,1776,1520,1478]
1804. 7 | 12 | 20 [avatar split clause 748,1776,1520,1483]
1805. 8 | 12 | 20 [avatar split clause 747,1776,1520,1488]
1806. 9 | 12 | 20 [avatar split clause 746,1776,1520,1493]
1807. ~1 | 13 | 20 [avatar split clause 745,1776,1531,1455]
1808. 4 | 13 | 20 [avatar split clause 744,1776,1531,1468]
1809. 5 | 13 | 20 [avatar split clause 743,1776,1531,1473]
1810. 6 | 13 | 20 [avatar split clause 742,1776,1531,1478]
1811. 7 | 13 | 20 [avatar split clause 741,1776,1531,1483]
1812. 8 | 13 | 20 [avatar split clause 740,1776,1531,1488]
1813. 9 | 13 | 20 [avatar split clause 739,1776,1531,1493]
1814. ~1 | 14 | 20 [avatar split clause 738,1776,1542,1455]
1815. 4 | 14 | 20 [avatar split clause 737,1776,1542,1468]
1816. 5 | 14 | 20 [avatar split clause 736,1776,1542,1473]
1817. 6 | 14 | 20 [avatar split clause 735,1776,1542,1478]
1818. 7 | 14 | 20 [avatar split clause 734,1776,1542,1483]
1819. 8 | 14 | 20 [avatar split clause 733,1776,1542,1488]
1820. 9 | 14 | 20 [avatar split clause 732,1776,1542,1493]
1821. ~1 | 15 | 20 [avatar split clause 731,1776,1553,1455]
1822. 4 | 15 | 20 [avatar split clause 730,1776,1553,1468]
1823. 5 | 15 | 20 [avatar split clause 729,1776,1553,1473]
1824. 6 | 15 | 20 [avatar split clause 728,1776,1553,1478]
1825. 7 | 15 | 20 [avatar split clause 727,1776,1553,1483]
1826. 8 | 15 | 20 [avatar split clause 726,1776,1553,1488]
1827. 9 | 15 | 20 [avatar split clause 725,1776,1553,1493]
1829. 21 <=> v1_relat_1(sK0) [avatar definition]
1832. ~1 | ~2 | 21 [avatar split clause 724,1829,1459,1455]
1833. 4 | ~2 | 21 [avatar split clause 723,1829,1459,1468]
1834. 5 | ~2 | 21 [avatar split clause 722,1829,1459,1473]
1835. 6 | ~2 | 21 [avatar split clause 721,1829,1459,1478]
1836. 7 | ~2 | 21 [avatar split clause 720,1829,1459,1483]
1837. 8 | ~2 | 21 [avatar split clause 719,1829,1459,1488]
1838. 9 | ~2 | 21 [avatar split clause 718,1829,1459,1493]
1839. ~1 | 10 | 21 [avatar split clause 717,1829,1498,1455]
1840. 4 | 10 | 21 [avatar split clause 716,1829,1498,1468]
1841. 5 | 10 | 21 [avatar split clause 715,1829,1498,1473]
1842. 6 | 10 | 21 [avatar split clause 714,1829,1498,1478]
1843. 7 | 10 | 21 [avatar split clause 713,1829,1498,1483]
1844. 8 | 10 | 21 [avatar split clause 712,1829,1498,1488]
1845. 9 | 10 | 21 [avatar split clause 711,1829,1498,1493]
1846. ~1 | 11 | 21 [avatar split clause 710,1829,1509,1455]
1847. 4 | 11 | 21 [avatar split clause 709,1829,1509,1468]
1848. 5 | 11 | 21 [avatar split clause 708,1829,1509,1473]
1849. 6 | 11 | 21 [avatar split clause 707,1829,1509,1478]
1850. 7 | 11 | 21 [avatar split clause 706,1829,1509,1483]
1851. 8 | 11 | 21 [avatar split clause 705,1829,1509,1488]
1852. 9 | 11 | 21 [avatar split clause 704,1829,1509,1493]
1853. ~1 | 12 | 21 [avatar split clause 703,1829,1520,1455]
1854. 4 | 12 | 21 [avatar split clause 702,1829,1520,1468]
1855. 5 | 12 | 21 [avatar split clause 701,1829,1520,1473]
1856. 6 | 12 | 21 [avatar split clause 700,1829,1520,1478]
1857. 7 | 12 | 21 [avatar split clause 699,1829,1520,1483]
1858. 8 | 12 | 21 [avatar split clause 698,1829,1520,1488]
1859. 9 | 12 | 21 [avatar split clause 697,1829,1520,1493]
1860. ~1 | 13 | 21 [avatar split clause 696,1829,1531,1455]
1861. 4 | 13 | 21 [avatar split clause 695,1829,1531,1468]
1862. 5 | 13 | 21 [avatar split clause 694,1829,1531,1473]
1863. 6 | 13 | 21 [avatar split clause 693,1829,1531,1478]
1864. 7 | 13 | 21 [avatar split clause 692,1829,1531,1483]
1865. 8 | 13 | 21 [avatar split clause 691,1829,1531,1488]
1866. 9 | 13 | 21 [avatar split clause 690,1829,1531,1493]
1867. ~1 | 14 | 21 [avatar split clause 689,1829,1542,1455]
1868. 4 | 14 | 21 [avatar split clause 688,1829,1542,1468]
1869. 5 | 14 | 21 [avatar split clause 687,1829,1542,1473]
1870. 6 | 14 | 21 [avatar split clause 686,1829,1542,1478]
1871. 7 | 14 | 21 [avatar split clause 685,1829,1542,1483]
1872. 8 | 14 | 21 [avatar split clause 684,1829,1542,1488]
1873. 9 | 14 | 21 [avatar split clause 683,1829,1542,1493]
1874. ~1 | 15 | 21 [avatar split clause 682,1829,1553,1455]
1875. 4 | 15 | 21 [avatar split clause 681,1829,1553,1468]
1876. 5 | 15 | 21 [avatar split clause 680,1829,1553,1473]
1877. 6 | 15 | 21 [avatar split clause 679,1829,1553,1478]
1878. 7 | 15 | 21 [avatar split clause 678,1829,1553,1483]
1879. 8 | 15 | 21 [avatar split clause 677,1829,1553,1488]
1880. 9 | 15 | 21 [avatar split clause 676,1829,1553,1493]
6154. k7_glib_003(sK0) != k7_glib_003(sK0) | ~v3_glib_003(sK0) | ~v1_glib_000(sK0) | ~v1_finset_1(sK0) | ~v1_funct_1(sK0) | ~v4_relat_1(sK0,k5_numbers) | ~v1_relat_1(sK0) <- (~3) [superposition 1465,1248] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6155. ~v3_glib_003(sK0) | ~v1_glib_000(sK0) | ~v1_finset_1(sK0) | ~v1_funct_1(sK0) | ~v4_relat_1(sK0,k5_numbers) | ~v1_relat_1(sK0) <- (~3) [trivial inequality removal 6154] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6165. ~v1_glib_000(sK0) | ~v1_finset_1(sK0) | ~v1_funct_1(sK0) | ~v4_relat_1(sK0,k5_numbers) | ~v1_relat_1(sK0) <- (~3, 16) [subsumption resolution 6155,1566] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6166. ~v1_finset_1(sK0) | ~v1_funct_1(sK0) | ~v4_relat_1(sK0,k5_numbers) | ~v1_relat_1(sK0) <- (~3, 16, 17) [subsumption resolution 6165,1619] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6167. ~v1_funct_1(sK0) | ~v4_relat_1(sK0,k5_numbers) | ~v1_relat_1(sK0) <- (~3, 16, 17, 18) [subsumption resolution 6166,1672] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6168. ~v4_relat_1(sK0,k5_numbers) | ~v1_relat_1(sK0) <- (~3, 16, 17, 18, 19) [subsumption resolution 6167,1725] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6169. ~v1_relat_1(sK0) <- (~3, 16, 17, 18, 19, 20) [subsumption resolution 6168,1778] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6172. ~21 | 3 | ~16 | ~17 | ~18 | ~19 | ~20 [avatar split clause 6169,1776,1723,1670,1617,1564,1463,1829]
6261. k6_glib_003(sK1) != k6_glib_003(sK1) | ~v2_glib_003(sK1) | ~v1_glib_000(sK1) | ~v1_finset_1(sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,k5_numbers) | ~v1_relat_1(sK1) <- (~2) [superposition 1461,1247] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6262. ~v2_glib_003(sK1) | ~v1_glib_000(sK1) | ~v1_finset_1(sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,k5_numbers) | ~v1_relat_1(sK1) <- (~2) [trivial inequality removal 6261] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6264. ~v1_glib_000(sK1) | ~v1_finset_1(sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,k5_numbers) | ~v1_relat_1(sK1) <- (~2, 10) [subsumption resolution 6262,1500] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6265. ~v1_finset_1(sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,k5_numbers) | ~v1_relat_1(sK1) <- (~2, 10, 11) [subsumption resolution 6264,1511] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6266. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,k5_numbers) | ~v1_relat_1(sK1) <- (~2, 10, 11, 12) [subsumption resolution 6265,1522] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6267. ~v4_relat_1(sK1,k5_numbers) | ~v1_relat_1(sK1) <- (~2, 10, 11, 12, 13) [subsumption resolution 6266,1533] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6268. ~v1_relat_1(sK1) <- (~2, 10, 11, 12, 13, 14) [subsumption resolution 6267,1544] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6271. ~15 | 2 | ~10 | ~11 | ~12 | ~13 | ~14 [avatar split clause 6268,1542,1531,1520,1509,1498,1459,1553]
6310. k5_glib_003(sK2) != k5_glib_003(sK2) | ~v1_glib_003(sK2) | ~v1_glib_000(sK2) | ~v1_finset_1(sK2) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,k5_numbers) | ~v1_relat_1(sK2) <- (~1) [superposition 1457,1249] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6311. ~v1_glib_003(sK2) | ~v1_glib_000(sK2) | ~v1_finset_1(sK2) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,k5_numbers) | ~v1_relat_1(sK2) <- (~1) [trivial inequality removal 6310] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6321. ~v1_glib_000(sK2) | ~v1_finset_1(sK2) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,k5_numbers) | ~v1_relat_1(sK2) <- (~1, 4) [subsumption resolution 6311,1470] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6322. ~v1_finset_1(sK2) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,k5_numbers) | ~v1_relat_1(sK2) <- (~1, 4, 5) [subsumption resolution 6321,1475] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6323. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,k5_numbers) | ~v1_relat_1(sK2) <- (~1, 4, 5, 6) [subsumption resolution 6322,1480] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6324. ~v4_relat_1(sK2,k5_numbers) | ~v1_relat_1(sK2) <- (~1, 4, 5, 6, 7) [subsumption resolution 6323,1485] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6325. ~v1_relat_1(sK2) <- (~1, 4, 5, 6, 7, 8) [subsumption resolution 6324,1490] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6326. $false <- (~1, 4, 5, 6, 7, 8, 9) [subsumption resolution 6325,1495] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
6327. 1 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 [avatar contradiction clause 6326]
6328. $false [avatar sat refutation 1466,1471,1476,1481,1486,1491,1496,1501,1502,1503,1504,1505,1506,1507,1512,1513,1514,1515,1516,1517,1518,1523,1524,1525,1526,1527,1528,1529,1534,1535,1536,1537,1538,1539,1540,1545,1546,1547,1548,1549,1550,1551,1556,1557,1558,1559,1560,1561,1562,1567,1568,1569,1570,1571,1572,1573,1574,1575,1576,1577,1578,1579,1580,1581,1582,1583,1584,1585,1586,1587,1588,1589,1590,1591,1592,1593,1594,1595,1596,1597,1598,1599,1600,1601,1602,1603,1604,1605,1606,1607,1608,1609,1610,1611,1612,1613,1614,1615,1620,1621,1622,1623,1624,1625,1626,1627,1628,1629,1630,1631,1632,1633,1634,1635,1636,1637,1638,1639,1640,1641,1642,1643,1644,1645,1646,1647,1648,1649,1650,1651,1652,1653,1654,1655,1656,1657,1658,1659,1660,1661,1662,1663,1664,1665,1666,1667,1668,1673,1674,1675,1676,1677,1678,1679,1680,1681,1682,1683,1684,1685,1686,1687,1688,1689,1690,1691,1692,1693,1694,1695,1696,1697,1698,1699,1700,1701,1702,1703,1704,1705,1706,1707,1708,1709,1710,1711,1712,1713,1714,1715,1716,1717,1718,1719,1720,1721,1726,1727,1728,1729,1730,1731,1732,1733,1734,1735,1736,1737,1738,1739,1740,1741,1742,1743,1744,1745,1746,1747,1748,1749,1750,1751,1752,1753,1754,1755,1756,1757,1758,1759,1760,1761,1762,1763,1764,1765,1766,1767,1768,1769,1770,1771,1772,1773,1774,1779,1780,1781,1782,1783,1784,1785,1786,1787,1788,1789,1790,1791,1792,1793,1794,1795,1796,1797,1798,1799,1800,1801,1802,1803,1804,1805,1806,1807,1808,1809,1810,1811,1812,1813,1814,1815,1816,1817,1818,1819,1820,1821,1822,1823,1824,1825,1826,1827,1832,1833,1834,1835,1836,1837,1838,1839,1840,1841,1842,1843,1844,1845,1846,1847,1848,1849,1850,1851,1852,1853,1854,1855,1856,1857,1858,1859,1860,1861,1862,1863,1864,1865,1866,1867,1868,1869,1870,1871,1872,1873,1874,1875,1876,1877,1878,1879,1880,6172,6271,6327] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:5616,thDist:-5616}
% SZS output end Proof for glib_003__t4_glib_003
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 264
% >>> Preprocessing
% Introduced skolems: 62
% Pure predicates: 15
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 778
% Generated clauses: 2767
% Active clauses: 529
% Passive clauses: 1035
% Final active clauses: 503
% Final passive clauses: 469
% Main loop iterations started: 560
% 
% >>> Simplifying Inferences
% Duplicate literals: 105
% Trivial inequalities: 3
% Fw subsumption resolutions: 453
% Fw demodulations: 14
% Bw demodulations: 8
% 
% >>> Deletion Inferences
% Simple tautologies: 69
% Forward subsumptions: 188
% Backward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 867
% Unit resulting resolution: 5
% Forward superposition: 13
% Backward superposition: 42
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 513
% Split components: 1439
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 970
% SAT solver unit clauses: 192
% SAT solver binary clauses: 188
% 
% Memory used [KB]: 13688
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% sat proof minimization: 0.001 s
% term sharing: 0.001 s
% congruence closure: 0.005 s
% backward subsumption index maintenance: 0.001 s
% backward superposition index maintenance: 0.001 s
% backward demodulation index maintenance: 0.001 s
% forward subsumption: 0.002 s
% resolution: 0.003 s ( own 0.002 s ) 
% unit resulting resolution: 0.003 s
% other: 0.027 s ( own 0.008 s ) 

