i: [1,52,0,2,1,1,0,0] t26_projpl_1
a: [3,104,0,2,1,52]
s: 104
i: [1,51,0,2,1,1,0,0] t26_projpl_1
a: [3,109,0,2,1,51]
s: 109
i: [1,57,0,4,1,1,0,0] t26_projpl_1
a: [3,79,0,4,1,57]
s: 79
i: [1,55,0,4,1,1,0,0] t26_projpl_1
a: [3,89,0,4,1,55]
s: 89
i: [1,54,0,4,1,1,0,0] t26_projpl_1
a: [3,94,0,4,1,54]
s: 94
i: [1,53,0,4,1,1,0,0] t26_projpl_1
a: [3,99,0,4,1,53]
s: 99
i: [1,56,0,5,1,1,0,0] t26_projpl_1
a: [3,84,0,5,1,56]
s: 84
i: [1,73,0,2,1,0,0,2] rc2_xboole_0
a: [3,114,0,2,1,73]
s: 114
i: [1,74,0,2,1,0,0,2] rc1_xboole_0
a: [3,119,0,2,1,74]
s: 119
i: [1,75,0,2,1,0,0,1] existence_l1_incsp_1
a: [3,124,0,2,1,75]
s: 124
i: [1,72,0,4,1,0,0,1] existence_m1_subset_1
s: 72
i: [1,58,0,5,2,0,0,2] dt_u2_incsp_1
s: 58
i: [1,59,0,5,2,0,0,1] dt_u1_incsp_1
s: 59
i: [1,63,0,9,3,0,0,1] d5_incproj
s: 63
i: [1,64,0,9,3,0,0,1] d5_incproj
s: 64
i: [1,71,0,13,4,0,0,255] d5_zfmisc_1
s: 71
d: [2,126,1,9,3,1,40,71,79]
a: [3,134,1,3,1,126]
s: 134
r: 84
r: 89
r: 79
d: [2,142,0,4,1,2,23,79,134]
a: [3,151,0,4,1,142]
s: 151
d: [2,140,0,5,1,2,23,84,134]
a: [3,146,0,5,1,140]
s: 146
i: [1,65,0,18,5,0,0,1] d5_incproj
s: 65
i: [1,60,0,19,5,0,0,1] d5_incproj
s: 60
i: [1,61,0,19,5,0,0,1] d5_incproj
s: 61
i: [1,62,0,19,5,0,0,1] d5_incproj
s: 62
i: [1,66,0,27,6,0,0,1] t5_projpl_1
s: 66
i: [1,67,0,27,6,0,0,1] t5_projpl_1
s: 67
i: [1,68,0,27,6,0,0,1] t5_projpl_1
s: 68
i: [1,69,0,27,6,0,0,1] t5_projpl_1
s: 69
i: [1,70,0,35,9,0,0,1] t5_projpl_1
s: 70
d: [2,168,1,30,8,1,40,70,146]
d: [2,169,1,22,6,1,18,168]
d: [2,170,1,18,5,2,21,169,99]
d: [2,171,1,14,4,3,21,170,94]
d: [2,172,1,12,3,4,21,171,104]
s: 172
d: [2,159,1,44,9,0,40,67,65]
d: [2,160,1,42,8,0,18,159]
d: [2,161,1,37,7,0,21,160,64]
s: 161
d: [2,162,1,44,9,0,40,68,65]
d: [2,163,1,42,8,0,18,162]
d: [2,164,1,37,7,0,21,163,64]
s: 164
d: [2,173,2,26,6,4,40,172,61]
d: [2,178,2,22,5,4,21,173,94]
d: [2,179,2,20,4,5,21,178,109]
d: [2,180,2,18,3,5,21,179,104]
s: 180
d: [2,165,1,44,9,0,40,69,65]
d: [2,166,1,42,8,0,18,165]
d: [2,167,1,37,7,0,21,166,64]
s: 167
d: [2,206,3,23,6,5,40,180,62]
d: [2,207,3,19,5,5,18,206]
d: [2,208,3,15,4,5,21,207,99]
d: [2,209,3,11,3,5,21,208,94]
d: [2,210,3,9,2,5,21,209,109]
d: [2,211,3,7,1,5,21,210,104]
a: [3,215,3,7,1,211]
s: 215
d: [2,225,4,0,0,5,52,104,109,94,99,215,60]
e: 225
d: [2,226,4,12,4,1,40,215,60]
d: [2,228,4,8,3,2,21,226,99]
d: [2,229,4,4,2,3,21,228,94]
d: [2,230,4,2,1,4,21,229,109]
d: [2,231,4,0,0,5,21,230,104]
e: 231
r: 134
r: 180
r: 172
r: 146
r: 215
s: 79
s: 89
a: [3,133,1,3,1,126]
s: 133
s: 84
a: [3,138,1,3,1,126]
s: 138
r: 84
r: 89
r: 79
d: [2,242,0,4,1,2,23,79,138]
a: [3,252,0,4,1,242]
s: 252
a: [3,145,0,5,1,140]
s: 145
a: [3,129,1,3,1,126]
s: 129
d: [2,240,0,5,1,2,23,84,138]
a: [3,247,0,5,1,240]
s: 247
a: [3,214,3,7,1,211]
s: 214
d: [2,257,1,30,8,1,40,247,70]
d: [2,258,1,22,6,1,18,257]
d: [2,259,1,18,5,2,21,258,94]
d: [2,260,1,14,4,3,21,259,99]
d: [2,261,1,12,3,4,21,260,104]
s: 261
d: [2,235,1,30,8,1,40,84,70]
d: [2,236,1,26,7,2,21,235,89]
d: [2,237,1,22,6,3,21,236,94]
d: [2,238,1,18,5,4,21,237,99]
d: [2,239,1,16,4,5,21,238,104]
s: 239
d: [2,262,2,26,6,4,40,261,61]
d: [2,267,2,22,5,4,21,262,94]
d: [2,268,2,20,4,5,21,267,109]
d: [2,269,2,18,3,5,21,268,104]
s: 269
d: [2,279,3,23,6,5,40,269,62]
d: [2,280,3,19,5,5,18,279]
d: [2,281,3,12,4,6,21,280,214]
d: [2,282,3,8,3,6,21,281,99]
d: [2,283,3,4,2,6,21,282,94]
d: [2,284,3,2,1,6,21,283,109]
d: [2,285,3,0,0,6,21,284,104]
e: 285
r: 129
r: 138
r: 269
r: 261
r: 247
s: 79
a: [3,130,1,3,1,126]
s: 130
r: 79
r: 94
r: 145
r: 151
r: 214
r: 252
s: 89
d: [2,290,0,4,1,2,23,79,130]
a: [3,299,0,4,1,290]
s: 299
a: [3,137,1,3,1,126]
s: 137
d: [2,293,0,4,1,2,23,151,130]
a: [3,309,0,4,1,293]
s: 309
s: 84
a: [3,246,0,5,1,240]
s: 246
d: [2,292,0,5,1,2,23,145,130]
a: [3,304,0,5,1,292]
s: 304
d: [2,322,0,5,1,2,22,84,130]
a: [3,326,0,5,1,322]
s: 326
d: [2,330,1,30,8,1,40,326,70]
d: [2,331,1,22,6,1,18,330]
d: [2,332,1,18,5,2,21,331,89]
d: [2,333,1,14,4,3,21,332,99]
d: [2,334,1,12,3,4,21,333,104]
s: 334
d: [2,294,3,7,1,2,23,214,130]
a: [3,314,3,7,1,294]
s: 314
s: 239
d: [2,335,2,26,6,4,40,334,61]
d: [2,340,2,22,5,4,21,335,89]
d: [2,341,2,20,4,5,21,340,109]
d: [2,342,2,18,3,5,21,341,104]
s: 342
d: [2,352,3,23,6,5,40,342,62]
d: [2,353,3,19,5,5,18,352]
d: [2,354,3,15,4,5,21,353,99]
d: [2,355,3,11,3,5,21,354,89]
d: [2,356,3,9,2,5,21,355,109]
d: [2,357,3,7,1,5,21,356,104]
a: [3,361,3,7,1,357]
s: 361
d: [2,363,4,0,0,5,52,104,109,89,99,361,60]
e: 363
d: [2,364,4,12,4,1,40,361,60]
d: [2,366,4,8,3,2,21,364,99]
d: [2,367,4,4,2,3,21,366,89]
d: [2,368,4,2,1,4,21,367,109]
d: [2,369,4,0,0,5,21,368,104]
e: 369
d: [2,287,1,9,3,1,40,79,71]
d: [2,288,1,6,2,2,21,287,133]
f: 369,357,322,288,285,240,231,211,140,51,52,53,54,55,56,57
% Refutation found. Thanks to Tanya!
% SZS status Theorem for projpl_1__t26_projpl_1
% SZS output start Proof for projpl_1__t26_projpl_1
1. ! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (~r4_projpl_1(X0,X1,X2,X3) => r1_zfmisc_1(X1,X2,X3)))))) [input t26_projpl_1]
2. ~! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (~r4_projpl_1(X0,X1,X2,X3) => r1_zfmisc_1(X1,X2,X3)))))) [negated conjecture 1]
3. ! [X0] : (l1_incsp_1(X0) => (v6_incsp_1(X0) <=> ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ? [X3] : (r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,X1,X3) & m1_subset_1(X3,u2_incsp_1(X0))))))) [input d5_incproj]
4. ! [X0] : ! [X1] : ! [X2] : (r1_zfmisc_1(X0,X1,X2) <=> (X1 != X2 & X0 != X2 & X0 != X1)) [input d5_zfmisc_1]
13. ! [X0] : (l1_incsp_1(X0) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (r4_projpl_1(X0,X1,X2,X3) <=> ? [X4] : (r1_incsp_1(X0,X3,X4) & r1_incsp_1(X0,X2,X4) & r1_incsp_1(X0,X1,X4) & m1_subset_1(X4,u2_incsp_1(X0)))))))) [input t5_projpl_1]
14. ! [X0,X1,X2] : (r1_zfmisc_1(X0,X1,X2) <=> (X1 != X2 & X0 != X2 & X0 != X1)) [flattening 4]
15. ! [X0,X1,X2] : ((X1 != X2 & X0 != X2 & X0 != X1) => r1_zfmisc_1(X0,X1,X2)) [unused predicate definition removal 14]
16. ~! [X0] : ((l1_incsp_1(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (~r4_projpl_1(X0,X1,X2,X3) => r1_zfmisc_1(X1,X2,X3)))))) [pure predicate removal 2]
17. ~! [X0] : ((l1_incsp_1(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (~r4_projpl_1(X0,X1,X2,X3) => r1_zfmisc_1(X1,X2,X3)))))) [pure predicate removal 16]
18. ~! [X0] : ((l1_incsp_1(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (~r4_projpl_1(X0,X1,X2,X3) => r1_zfmisc_1(X1,X2,X3)))))) [pure predicate removal 17]
19. ~! [X0] : ((l1_incsp_1(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => (~r4_projpl_1(X0,X1,X2,X3) => r1_zfmisc_1(X1,X2,X3)))))) [pure predicate removal 18]
20. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_zfmisc_1(X1,X2,X3) & ~r4_projpl_1(X0,X1,X2,X3)) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & (l1_incsp_1(X0) & v6_incsp_1(X0))) [ennf transformation 19]
21. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_zfmisc_1(X1,X2,X3) & ~r4_projpl_1(X0,X1,X2,X3) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v6_incsp_1(X0)) [flattening 20]
24. ! [X0] : ((v6_incsp_1(X0) <=> ! [X1] : (! [X2] : (? [X3] : (r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,X1,X3) & m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0)))) | ~l1_incsp_1(X0)) [ennf transformation 3]
25. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r4_projpl_1(X0,X1,X2,X3) <=> ? [X4] : (r1_incsp_1(X0,X3,X4) & r1_incsp_1(X0,X2,X4) & r1_incsp_1(X0,X1,X4) & m1_subset_1(X4,u2_incsp_1(X0)))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [ennf transformation 13]
26. ! [X0,X1,X2] : (r1_zfmisc_1(X0,X1,X2) | (X1 = X2 | X0 = X2 | X0 = X1)) [ennf transformation 15]
27. ! [X0,X1,X2] : (r1_zfmisc_1(X0,X1,X2) | X1 = X2 | X0 = X2 | X0 = X1) [flattening 26]
28. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_zfmisc_1(X1,X2,X3) & ~r4_projpl_1(X0,X1,X2,X3) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v6_incsp_1(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r1_zfmisc_1(X1,X2,X3) & ~r4_projpl_1(sK0,X1,X2,X3) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v6_incsp_1(sK0)) [choice axiom]
29. ? [X1] : (? [X2] : (? [X3] : (~r1_zfmisc_1(X1,X2,X3) & ~r4_projpl_1(sK0,X1,X2,X3) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) => (? [X2] : (? [X3] : (~r1_zfmisc_1(sK1,X2,X3) & ~r4_projpl_1(sK0,sK1,X2,X3) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) [choice axiom]
30. ? [X2] : (? [X3] : (~r1_zfmisc_1(sK1,X2,X3) & ~r4_projpl_1(sK0,sK1,X2,X3) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) => (? [X3] : (~r1_zfmisc_1(sK1,sK2,X3) & ~r4_projpl_1(sK0,sK1,sK2,X3) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) [choice axiom]
31. ? [X3] : (~r1_zfmisc_1(sK1,sK2,X3) & ~r4_projpl_1(sK0,sK1,sK2,X3) & m1_subset_1(X3,u1_incsp_1(sK0))) => (~r1_zfmisc_1(sK1,sK2,sK3) & ~r4_projpl_1(sK0,sK1,sK2,sK3) & m1_subset_1(sK3,u1_incsp_1(sK0))) [choice axiom]
32. (((~r1_zfmisc_1(sK1,sK2,sK3) & ~r4_projpl_1(sK0,sK1,sK2,sK3) & m1_subset_1(sK3,u1_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v6_incsp_1(sK0) [skolemisation 21,31,30,29,28]
33. ! [X0] : (((v6_incsp_1(X0) | ? [X1] : (? [X2] : (! [X3] : (~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3) | ~m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0)))) & (! [X1] : (! [X2] : (? [X3] : (r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,X1,X3) & m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~v6_incsp_1(X0))) | ~l1_incsp_1(X0)) [nnf transformation 24]
34. ! [X0] : (((v6_incsp_1(X0) | ? [X1] : (? [X2] : (! [X3] : (~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3) | ~m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0)))) & (! [X4] : (! [X5] : (? [X6] : (r1_incsp_1(X0,X5,X6) & r1_incsp_1(X0,X4,X6) & m1_subset_1(X6,u2_incsp_1(X0))) | ~m1_subset_1(X5,u1_incsp_1(X0))) | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~v6_incsp_1(X0))) | ~l1_incsp_1(X0)) [rectify 33]
35. ! [X0] : (? [X1] : (? [X2] : (! [X3] : (~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3) | ~m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) => (? [X2] : (! [X3] : (~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,sK4(X0),X3) | ~m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(sK4(X0),u1_incsp_1(X0)))) [choice axiom]
36. ! [X0] : (? [X2] : (! [X3] : (~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,sK4(X0),X3) | ~m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) => (! [X3] : (~r1_incsp_1(X0,sK5(X0),X3) | ~r1_incsp_1(X0,sK4(X0),X3) | ~m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(sK5(X0),u1_incsp_1(X0)))) [choice axiom]
37. ! [X5,X4,X0] : (? [X6] : (r1_incsp_1(X0,X5,X6) & r1_incsp_1(X0,X4,X6) & m1_subset_1(X6,u2_incsp_1(X0))) => (r1_incsp_1(X0,X5,sK6(X0,X4,X5)) & r1_incsp_1(X0,X4,sK6(X0,X4,X5)) & m1_subset_1(sK6(X0,X4,X5),u2_incsp_1(X0)))) [choice axiom]
38. ! [X0] : (((v6_incsp_1(X0) | ((! [X3] : (~r1_incsp_1(X0,sK5(X0),X3) | ~r1_incsp_1(X0,sK4(X0),X3) | ~m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(sK5(X0),u1_incsp_1(X0))) & m1_subset_1(sK4(X0),u1_incsp_1(X0)))) & (! [X4] : (! [X5] : ((r1_incsp_1(X0,X5,sK6(X0,X4,X5)) & r1_incsp_1(X0,X4,sK6(X0,X4,X5)) & m1_subset_1(sK6(X0,X4,X5),u2_incsp_1(X0))) | ~m1_subset_1(X5,u1_incsp_1(X0))) | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~v6_incsp_1(X0))) | ~l1_incsp_1(X0)) [skolemisation 34,37,36,35]
39. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r4_projpl_1(X0,X1,X2,X3) | ! [X4] : (~r1_incsp_1(X0,X3,X4) | ~r1_incsp_1(X0,X2,X4) | ~r1_incsp_1(X0,X1,X4) | ~m1_subset_1(X4,u2_incsp_1(X0)))) & (? [X4] : (r1_incsp_1(X0,X3,X4) & r1_incsp_1(X0,X2,X4) & r1_incsp_1(X0,X1,X4) & m1_subset_1(X4,u2_incsp_1(X0))) | ~r4_projpl_1(X0,X1,X2,X3))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [nnf transformation 25]
40. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r4_projpl_1(X0,X1,X2,X3) | ! [X4] : (~r1_incsp_1(X0,X3,X4) | ~r1_incsp_1(X0,X2,X4) | ~r1_incsp_1(X0,X1,X4) | ~m1_subset_1(X4,u2_incsp_1(X0)))) & (? [X5] : (r1_incsp_1(X0,X3,X5) & r1_incsp_1(X0,X2,X5) & r1_incsp_1(X0,X1,X5) & m1_subset_1(X5,u2_incsp_1(X0))) | ~r4_projpl_1(X0,X1,X2,X3))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [rectify 39]
41. ! [X3,X2,X1,X0] : (? [X5] : (r1_incsp_1(X0,X3,X5) & r1_incsp_1(X0,X2,X5) & r1_incsp_1(X0,X1,X5) & m1_subset_1(X5,u2_incsp_1(X0))) => (r1_incsp_1(X0,X3,sK7(X0,X1,X2,X3)) & r1_incsp_1(X0,X2,sK7(X0,X1,X2,X3)) & r1_incsp_1(X0,X1,sK7(X0,X1,X2,X3)) & m1_subset_1(sK7(X0,X1,X2,X3),u2_incsp_1(X0)))) [choice axiom]
42. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r4_projpl_1(X0,X1,X2,X3) | ! [X4] : (~r1_incsp_1(X0,X3,X4) | ~r1_incsp_1(X0,X2,X4) | ~r1_incsp_1(X0,X1,X4) | ~m1_subset_1(X4,u2_incsp_1(X0)))) & ((r1_incsp_1(X0,X3,sK7(X0,X1,X2,X3)) & r1_incsp_1(X0,X2,sK7(X0,X1,X2,X3)) & r1_incsp_1(X0,X1,sK7(X0,X1,X2,X3)) & m1_subset_1(sK7(X0,X1,X2,X3),u2_incsp_1(X0))) | ~r4_projpl_1(X0,X1,X2,X3))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0)) [skolemisation 40,41]
51. v6_incsp_1(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
52. l1_incsp_1(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
53. m1_subset_1(sK1,u1_incsp_1(sK0)) [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
54. m1_subset_1(sK2,u1_incsp_1(sK0)) [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
55. m1_subset_1(sK3,u1_incsp_1(sK0)) [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
56. ~r4_projpl_1(sK0,sK1,sK2,sK3) [cnf transformation 32] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
57. ~r1_zfmisc_1(sK1,sK2,sK3) [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
60. m1_subset_1(sK6(X0,X4,X5),u2_incsp_1(X0)) | ~m1_subset_1(X5,u1_incsp_1(X0)) | ~m1_subset_1(X4,u1_incsp_1(X0)) | ~v6_incsp_1(X0) | ~l1_incsp_1(X0) [cnf transformation 38] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
61. r1_incsp_1(X0,X4,sK6(X0,X4,X5)) | ~m1_subset_1(X5,u1_incsp_1(X0)) | ~m1_subset_1(X4,u1_incsp_1(X0)) | ~v6_incsp_1(X0) | ~l1_incsp_1(X0) [cnf transformation 38] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
62. r1_incsp_1(X0,X5,sK6(X0,X4,X5)) | ~m1_subset_1(X5,u1_incsp_1(X0)) | ~m1_subset_1(X4,u1_incsp_1(X0)) | ~v6_incsp_1(X0) | ~l1_incsp_1(X0) [cnf transformation 38] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
70. r4_projpl_1(X0,X1,X2,X3) | ~r1_incsp_1(X0,X3,X4) | ~r1_incsp_1(X0,X2,X4) | ~r1_incsp_1(X0,X1,X4) | ~m1_subset_1(X4,u2_incsp_1(X0)) | ~m1_subset_1(X3,u1_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) [cnf transformation 42] {a:0,w:35,wCS:14000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
71. r1_zfmisc_1(X0,X1,X2) | X1 = X2 | X0 = X2 | X0 = X1 [cnf transformation 27] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
77. 1 <=> r1_zfmisc_1(sK1,sK2,sK3) [avatar definition]
79. ~r1_zfmisc_1(sK1,sK2,sK3) <- (~1) [avatar component clause 77] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
80. ~1 [avatar split clause 57,77]
82. 2 <=> r4_projpl_1(sK0,sK1,sK2,sK3) [avatar definition]
84. ~r4_projpl_1(sK0,sK1,sK2,sK3) <- (~2) [avatar component clause 82] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
85. ~2 [avatar split clause 56,82]
87. 3 <=> m1_subset_1(sK3,u1_incsp_1(sK0)) [avatar definition]
89. m1_subset_1(sK3,u1_incsp_1(sK0)) <- (3) [avatar component clause 87] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
90. 3 [avatar split clause 55,87]
92. 4 <=> m1_subset_1(sK2,u1_incsp_1(sK0)) [avatar definition]
94. m1_subset_1(sK2,u1_incsp_1(sK0)) <- (4) [avatar component clause 92] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
95. 4 [avatar split clause 54,92]
97. 5 <=> m1_subset_1(sK1,u1_incsp_1(sK0)) [avatar definition]
99. m1_subset_1(sK1,u1_incsp_1(sK0)) <- (5) [avatar component clause 97] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
100. 5 [avatar split clause 53,97]
102. 6 <=> l1_incsp_1(sK0) [avatar definition]
104. l1_incsp_1(sK0) <- (6) [avatar component clause 102] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
105. 6 [avatar split clause 52,102]
107. 7 <=> v6_incsp_1(sK0) [avatar definition]
109. v6_incsp_1(sK0) <- (7) [avatar component clause 107] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
110. 7 [avatar split clause 51,107]
128. 11 <=> sK1 = sK2 [avatar definition]
130. sK1 = sK2 <- (11) [avatar component clause 128] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
132. 12 <=> sK1 = sK3 [avatar definition]
133. sK1 != sK3 <- (~12) [avatar component clause 132] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
134. sK1 = sK3 <- (12) [avatar component clause 132] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
136. 13 <=> sK2 = sK3 [avatar definition]
138. sK2 = sK3 <- (13) [avatar component clause 136] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
140. ~r4_projpl_1(sK0,sK1,sK2,sK1) <- (~2, 12) [backward demodulation 84,134] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
144. 14 <=> r4_projpl_1(sK0,sK1,sK2,sK1) [avatar definition]
146. ~r4_projpl_1(sK0,sK1,sK2,sK1) <- (~14) [avatar component clause 144] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
147. ~14 | 2 | ~12 [avatar split clause 140,132,82,144]
168. ~r1_incsp_1(sK0,sK1,X0) | ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (~14) [resolution 70,146] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
169. ~r1_incsp_1(sK0,sK1,X0) | ~r1_incsp_1(sK0,sK2,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (~14) [duplicate literal removal 168] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
170. ~r1_incsp_1(sK0,sK1,X0) | ~r1_incsp_1(sK0,sK2,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (5, ~14) [subsumption resolution 169,99] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
171. ~r1_incsp_1(sK0,sK1,X0) | ~r1_incsp_1(sK0,sK2,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 5, ~14) [subsumption resolution 170,94] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
172. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) <- (4, 5, 6, ~14) [subsumption resolution 171,104] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
173. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK2,X0)) | ~m1_subset_1(sK6(sK0,sK2,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, ~14) [resolution 172,61] {a:2,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
178. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK2,X0)) | ~m1_subset_1(sK6(sK0,sK2,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, ~14) [subsumption resolution 173,94] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
179. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK2,X0)) | ~m1_subset_1(sK6(sK0,sK2,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, ~14) [subsumption resolution 178,109] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
180. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK2,X0)) | ~m1_subset_1(sK6(sK0,sK2,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) <- (4, 5, 6, 7, ~14) [subsumption resolution 179,104] {a:2,w:18,wCS:1800,nSel:1,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
206. ~m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, ~14) [resolution 180,62] {a:3,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
207. ~m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, ~14) [duplicate literal removal 206] {a:3,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
208. ~m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, ~14) [subsumption resolution 207,99] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
209. ~m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, ~14) [subsumption resolution 208,94] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
210. ~m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, ~14) [subsumption resolution 209,109] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
211. ~m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) <- (4, 5, 6, 7, ~14) [subsumption resolution 210,104] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
213. 16 <=> m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) [avatar definition]
214. m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) <- (16) [avatar component clause 213] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:15,thDist:-15}
215. ~m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) <- (~16) [avatar component clause 213] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:15,thDist:-15}
216. ~16 | ~4 | ~5 | ~6 | ~7 | 14 [avatar split clause 211,144,107,102,97,92,213]
226. ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (~16) [resolution 215,60] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
228. ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (5, ~16) [subsumption resolution 226,99] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
229. ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, ~16) [subsumption resolution 228,94] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
230. ~l1_incsp_1(sK0) <- (4, 5, 7, ~16) [subsumption resolution 229,109] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
231. $false <- (4, 5, 6, 7, ~16) [subsumption resolution 230,104] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
232. ~4 | ~5 | ~6 | ~7 | 16 [avatar contradiction clause 231]
240. ~r4_projpl_1(sK0,sK1,sK2,sK2) <- (~2, 13) [backward demodulation 84,138] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
245. 17 <=> r4_projpl_1(sK0,sK1,sK2,sK2) [avatar definition]
247. ~r4_projpl_1(sK0,sK1,sK2,sK2) <- (~17) [avatar component clause 245] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
248. ~17 | 2 | ~13 [avatar split clause 240,136,82,245]
257. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (~17) [resolution 247,70] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
258. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (~17) [duplicate literal removal 257] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
259. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, ~17) [subsumption resolution 258,94] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
260. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 5, ~17) [subsumption resolution 259,99] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
261. ~r1_incsp_1(sK0,sK2,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) <- (4, 5, 6, ~17) [subsumption resolution 260,104] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
262. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK2,X0)) | ~m1_subset_1(sK6(sK0,sK2,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, ~17) [resolution 261,61] {a:2,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
267. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK2,X0)) | ~m1_subset_1(sK6(sK0,sK2,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, ~17) [subsumption resolution 262,94] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
268. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK2,X0)) | ~m1_subset_1(sK6(sK0,sK2,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, ~17) [subsumption resolution 267,109] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
269. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK2,X0)) | ~m1_subset_1(sK6(sK0,sK2,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) <- (4, 5, 6, 7, ~17) [subsumption resolution 268,104] {a:2,w:18,wCS:1800,nSel:1,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
279. ~m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, ~17) [resolution 269,62] {a:3,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
280. ~m1_subset_1(sK6(sK0,sK2,sK1),u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, ~17) [duplicate literal removal 279] {a:3,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
281. ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, 16, ~17) [subsumption resolution 280,214] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
282. ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, 16, ~17) [subsumption resolution 281,99] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
283. ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (4, 5, 6, 7, 16, ~17) [subsumption resolution 282,94] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
284. ~l1_incsp_1(sK0) <- (4, 5, 6, 7, 16, ~17) [subsumption resolution 283,109] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
285. $false <- (4, 5, 6, 7, 16, ~17) [subsumption resolution 284,104] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
286. ~4 | ~5 | ~6 | ~7 | ~16 | 17 [avatar contradiction clause 285]
287. sK2 = sK3 | sK1 = sK3 | sK1 = sK2 <- (~1) [resolution 79,71] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
288. sK2 = sK3 | sK1 = sK2 <- (~1, ~12) [subsumption resolution 287,133] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
289. 11 | 13 | 1 | 12 [avatar split clause 288,132,77,136,128]
322. ~r4_projpl_1(sK0,sK1,sK1,sK3) <- (~2, 11) [forward demodulation 84,130] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
324. 23 <=> r4_projpl_1(sK0,sK1,sK1,sK3) [avatar definition]
326. ~r4_projpl_1(sK0,sK1,sK1,sK3) <- (~23) [avatar component clause 324] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
327. ~23 | 2 | ~11 [avatar split clause 322,128,82,324]
330. ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (~23) [resolution 326,70] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
331. ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (~23) [duplicate literal removal 330] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
332. ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, ~23) [subsumption resolution 331,89] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
333. ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, 5, ~23) [subsumption resolution 332,99] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
334. ~r1_incsp_1(sK0,sK3,X0) | ~r1_incsp_1(sK0,sK1,X0) | ~m1_subset_1(X0,u2_incsp_1(sK0)) <- (3, 5, 6, ~23) [subsumption resolution 333,104] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
335. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK3,X0)) | ~m1_subset_1(sK6(sK0,sK3,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (3, 5, 6, ~23) [resolution 334,61] {a:2,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
340. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK3,X0)) | ~m1_subset_1(sK6(sK0,sK3,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (3, 5, 6, ~23) [subsumption resolution 335,89] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
341. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK3,X0)) | ~m1_subset_1(sK6(sK0,sK3,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, 5, 6, 7, ~23) [subsumption resolution 340,109] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
342. ~r1_incsp_1(sK0,sK1,sK6(sK0,sK3,X0)) | ~m1_subset_1(sK6(sK0,sK3,X0),u2_incsp_1(sK0)) | ~m1_subset_1(X0,u1_incsp_1(sK0)) <- (3, 5, 6, 7, ~23) [subsumption resolution 341,104] {a:2,w:18,wCS:1800,nSel:1,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
352. ~m1_subset_1(sK6(sK0,sK3,sK1),u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (3, 5, 6, 7, ~23) [resolution 342,62] {a:3,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
353. ~m1_subset_1(sK6(sK0,sK3,sK1),u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (3, 5, 6, 7, ~23) [duplicate literal removal 352] {a:3,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
354. ~m1_subset_1(sK6(sK0,sK3,sK1),u2_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (3, 5, 6, 7, ~23) [subsumption resolution 353,99] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
355. ~m1_subset_1(sK6(sK0,sK3,sK1),u2_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (3, 5, 6, 7, ~23) [subsumption resolution 354,89] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
356. ~m1_subset_1(sK6(sK0,sK3,sK1),u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) <- (3, 5, 6, 7, ~23) [subsumption resolution 355,109] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
357. ~m1_subset_1(sK6(sK0,sK3,sK1),u2_incsp_1(sK0)) <- (3, 5, 6, 7, ~23) [subsumption resolution 356,104] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
359. 24 <=> m1_subset_1(sK6(sK0,sK3,sK1),u2_incsp_1(sK0)) [avatar definition]
361. ~m1_subset_1(sK6(sK0,sK3,sK1),u2_incsp_1(sK0)) <- (~24) [avatar component clause 359] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:15,thDist:-15}
362. ~24 | ~3 | ~5 | ~6 | ~7 | 23 [avatar split clause 357,324,107,102,97,87,359]
364. ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (~24) [resolution 361,60] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
366. ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (5, ~24) [subsumption resolution 364,99] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
367. ~v6_incsp_1(sK0) | ~l1_incsp_1(sK0) <- (3, 5, ~24) [subsumption resolution 366,89] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
368. ~l1_incsp_1(sK0) <- (3, 5, 7, ~24) [subsumption resolution 367,109] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
369. $false <- (3, 5, 6, 7, ~24) [subsumption resolution 368,104] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
370. ~3 | ~5 | ~6 | ~7 | 24 [avatar contradiction clause 369]
371. $false [avatar sat refutation 80,85,90,95,100,105,110,147,216,232,248,286,289,327,362,370] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:367,thDist:-367}
% SZS output end Proof for projpl_1__t26_projpl_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% >>> Preprocessing
% Introduced skolems: 12
% Pure predicates: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 25
% Generated clauses: 229
% Active clauses: 57
% Passive clauses: 80
% Final active clauses: 35
% Final passive clauses: 10
% Main loop iterations started: 62
% 
% >>> Simplifying Inferences
% Duplicate literals: 43
% Fw subsumption resolutions: 86
% Fw demodulations: 3
% Bw demodulations: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 4
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 42
% Unit resulting resolution: 2
% Backward superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 8
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 34
% SAT solver unit clauses: 10
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.003 s

