i: [1,103,0,2,1,1,0,0] t7_incproj
a: [3,151,0,2,1,103]
s: 151
i: [1,102,0,2,1,1,0,0] t7_incproj
a: [3,156,0,2,1,102]
s: 156
i: [1,101,0,2,1,1,0,0] t7_incproj
a: [3,161,0,2,1,101]
s: 161
i: [1,100,0,2,1,1,0,0] t7_incproj
a: [3,166,0,2,1,100]
s: 166
i: [1,99,0,2,1,1,0,0] t7_incproj
a: [3,171,0,2,1,99]
s: 171
i: [1,104,0,4,1,1,0,0] t7_incproj
a: [3,146,0,4,1,104]
s: 146
i: [1,105,0,7,2,1,0,0] t7_incproj
s: 105
i: [1,106,0,2,1,0,0,255] spc1_boole
a: [3,176,0,2,1,106]
s: 176
i: [1,141,0,2,1,0,0,1] existence_l1_collsp
a: [3,181,0,2,1,141]
s: 181
i: [1,142,0,2,1,0,0,2] existence_l1_struct_0
a: [3,186,0,2,1,142]
s: 186
i: [1,107,0,4,2,0,0,1] dt_l1_collsp
s: 107
d: [2,189,1,2,1,1,40,107,151]
a: [3,194,1,2,1,189]
s: 194
i: [1,139,0,4,1,0,0,1] existence_m1_subset_1
s: 139
d: [2,201,1,5,1,0,40,139,105]
a: [3,205,1,5,1,201]
s: 205
i: [1,108,0,5,2,0,0,255] t6_boole
s: 108
i: [1,138,0,5,2,0,0,255] t7_boole
s: 138
d: [2,190,1,2,1,1,40,107,181]
a: [3,199,1,2,1,190]
s: 199
i: [1,109,0,6,3,0,0,255] cc7_struct_0
s: 109
i: [1,110,0,6,3,0,0,255] cc5_struct_0
s: 110
i: [1,112,0,6,3,0,0,255] cc2_struct_0
s: 112
i: [1,113,0,6,3,0,0,255] cc17_struct_0
s: 113
d: [2,209,1,4,2,1,40,110,186]
a: [3,215,1,2,1,209]
s: 215
i: [1,114,0,6,3,0,0,255] cc6_struct_0
s: 114
d: [2,211,2,4,2,1,40,110,199]
a: [3,224,2,2,1,211]
s: 224
i: [1,115,0,6,3,0,0,255] cc1_struct_0
s: 115
i: [1,118,0,6,3,0,0,255] cc4_struct_0
s: 118
i: [1,135,0,6,2,0,0,255] t1_subset
s: 135
i: [1,140,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 140
i: [1,116,0,7,3,0,0,255] cc8_struct_0
s: 116
i: [1,119,0,7,3,0,0,255] cc9_struct_0
s: 119
i: [1,120,0,7,3,0,0,255] cc11_struct_0
s: 120
i: [1,121,0,7,3,0,0,255] cc11_struct_0
s: 121
i: [1,123,0,7,3,0,0,255] fc9_struct_0
s: 123
i: [1,124,0,7,3,0,0,255] fc6_struct_0
s: 124
i: [1,131,0,7,3,0,0,2] fc2_struct_0
s: 131
i: [1,132,0,7,3,0,0,255] fc8_struct_0
s: 132
i: [1,133,0,7,3,0,0,255] fc7_struct_0
s: 133
i: [1,134,0,7,3,0,0,2] fc1_struct_0
s: 134
i: [1,137,0,7,3,0,0,1] t8_boole
s: 137
i: [1,136,0,8,3,0,0,255] t2_subset
s: 136
i: [1,122,0,9,4,0,0,255] cc10_struct_0
s: 122
d: [2,253,1,7,2,1,40,136,146]
a: [3,258,1,4,1,253]
s: 258
i: [1,125,0,15,6,0,0,1] t6_incproj
s: 125
a: [3,261,1,3,1,253]
s: 261
i: [1,126,0,15,6,0,0,1] t6_incproj
s: 126
i: [1,127,0,15,6,0,0,1] t6_incproj
s: 127
d: [2,277,1,14,6,0,40,125,105]
d: [2,279,1,12,5,1,21,277,151]
d: [2,280,1,10,4,2,21,279,156]
d: [2,281,1,8,3,3,21,280,161]
d: [2,282,1,6,2,4,21,281,166]
d: [2,283,1,4,1,5,21,282,171]
a: [3,287,1,4,1,283]
s: 287
i: [1,128,0,15,6,0,0,1] t6_incproj
s: 128
d: [2,290,1,14,6,0,40,126,105]
d: [2,292,1,12,5,1,21,290,151]
d: [2,293,1,10,4,2,21,292,156]
d: [2,294,1,8,3,3,21,293,161]
d: [2,295,1,6,2,4,21,294,166]
d: [2,296,1,4,1,5,21,295,171]
a: [3,300,1,4,1,296]
s: 300
d: [2,315,2,14,6,1,44,128,300]
d: [2,317,2,10,5,2,21,315,287]
d: [2,318,2,8,4,3,21,317,151]
d: [2,319,2,6,3,4,21,318,156]
d: [2,320,2,4,2,5,21,319,161]
d: [2,321,2,2,1,6,21,320,166]
d: [2,322,2,0,0,7,21,321,171]
e: 322
f: 322,296,283,99,100,101,102,103
% Refutation found. Thanks to Tanya!
% SZS status Theorem for incproj__t7_incproj
% SZS output start Proof for incproj__t7_incproj
1. ! [X0] : ((l1_collsp(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ~! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => X1 = X2))) [input t7_incproj]
2. ~! [X0] : ((l1_collsp(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ~! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => X1 = X2))) [negated conjecture 1]
33. ! [X0] : ((l1_collsp(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => ? [X1] : (? [X2] : (? [X3] : (X1 != X3 & X2 != X3 & X1 != X2 & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0)))) [input t6_incproj]
40. ? [X0] : (? [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_collsp(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
41. ? [X0] : (? [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_collsp(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) [flattening 40]
70. ! [X0] : (? [X1] : (? [X2] : (? [X3] : (X1 != X3 & X2 != X3 & X1 != X2 & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | (~l1_collsp(X0) | ~v4_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0))) [ennf transformation 33]
71. ! [X0] : (? [X1] : (? [X2] : (? [X3] : (X1 != X3 & X2 != X3 & X1 != X2 & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | ~l1_collsp(X0) | ~v4_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0)) [flattening 70]
86. ? [X0] : (? [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_collsp(X0) & v4_collsp(X0) & v3_collsp(X0) & v2_collsp(X0) & ~v2_struct_0(X0)) => (? [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_collsp(sK0) & v4_collsp(sK0) & v3_collsp(sK0) & v2_collsp(sK0) & ~v2_struct_0(sK0)) [choice axiom]
87. ? [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (! [X2] : (sK1 = X2 | ~m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
88. (! [X2] : (sK1 = X2 | ~m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_collsp(sK0) & v4_collsp(sK0) & v3_collsp(sK0) & v2_collsp(sK0) & ~v2_struct_0(sK0) [skolemisation 41,87,86]
89. ! [X0] : (? [X1] : (? [X2] : (? [X3] : (X1 != X3 & X2 != X3 & X1 != X2 & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) => (? [X2] : (? [X3] : (sK2(X0) != X3 & X2 != X3 & sK2(X0) != X2 & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(sK2(X0),u1_struct_0(X0)))) [choice axiom]
90. ! [X0] : (? [X2] : (? [X3] : (sK2(X0) != X3 & X2 != X3 & sK2(X0) != X2 & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) => (? [X3] : (sK2(X0) != X3 & sK3(X0) != X3 & sK2(X0) != sK3(X0) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(sK3(X0),u1_struct_0(X0)))) [choice axiom]
91. ! [X0] : (? [X3] : (sK2(X0) != X3 & sK3(X0) != X3 & sK2(X0) != sK3(X0) & m1_subset_1(X3,u1_struct_0(X0))) => (sK2(X0) != sK4(X0) & sK3(X0) != sK4(X0) & sK2(X0) != sK3(X0) & m1_subset_1(sK4(X0),u1_struct_0(X0)))) [choice axiom]
92. ! [X0] : ((((sK2(X0) != sK4(X0) & sK3(X0) != sK4(X0) & sK2(X0) != sK3(X0) & m1_subset_1(sK4(X0),u1_struct_0(X0))) & m1_subset_1(sK3(X0),u1_struct_0(X0))) & m1_subset_1(sK2(X0),u1_struct_0(X0))) | ~l1_collsp(X0) | ~v4_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0)) [skolemisation 71,91,90,89]
99. ~v2_struct_0(sK0) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. v2_collsp(sK0) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
101. v3_collsp(sK0) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
102. v4_collsp(sK0) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
103. l1_collsp(sK0) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
105. ~m1_subset_1(X2,u1_struct_0(sK0)) | sK1 = X2 [cnf transformation 88] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
125. m1_subset_1(sK2(X0),u1_struct_0(X0)) | ~l1_collsp(X0) | ~v4_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0) [cnf transformation 92] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
126. m1_subset_1(sK3(X0),u1_struct_0(X0)) | ~l1_collsp(X0) | ~v4_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0) [cnf transformation 92] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
128. sK2(X0) != sK3(X0) | ~l1_collsp(X0) | ~v4_collsp(X0) | ~v3_collsp(X0) | ~v2_collsp(X0) | v2_struct_0(X0) [cnf transformation 92] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
149. 2 <=> l1_collsp(sK0) [avatar definition]
151. l1_collsp(sK0) <- (2) [avatar component clause 149] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
152. 2 [avatar split clause 103,149]
154. 3 <=> v4_collsp(sK0) [avatar definition]
156. v4_collsp(sK0) <- (3) [avatar component clause 154] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
157. 3 [avatar split clause 102,154]
159. 4 <=> v3_collsp(sK0) [avatar definition]
161. v3_collsp(sK0) <- (4) [avatar component clause 159] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
162. 4 [avatar split clause 101,159]
164. 5 <=> v2_collsp(sK0) [avatar definition]
166. v2_collsp(sK0) <- (5) [avatar component clause 164] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
167. 5 [avatar split clause 100,164]
169. 6 <=> v2_struct_0(sK0) [avatar definition]
171. ~v2_struct_0(sK0) <- (~6) [avatar component clause 169] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
172. ~6 [avatar split clause 99,169]
277. ~l1_collsp(sK0) | ~v4_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) | sK1 = sK2(sK0) [resolution 125,105] {a:1,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
279. ~v4_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) | sK1 = sK2(sK0) <- (2) [subsumption resolution 277,151] {a:1,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
280. ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) | sK1 = sK2(sK0) <- (2, 3) [subsumption resolution 279,156] {a:1,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
281. ~v2_collsp(sK0) | v2_struct_0(sK0) | sK1 = sK2(sK0) <- (2, 3, 4) [subsumption resolution 280,161] {a:1,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
282. v2_struct_0(sK0) | sK1 = sK2(sK0) <- (2, 3, 4, 5) [subsumption resolution 281,166] {a:1,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
283. sK1 = sK2(sK0) <- (2, 3, 4, 5, ~6) [subsumption resolution 282,171] {a:1,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
285. 22 <=> sK1 = sK2(sK0) [avatar definition]
287. sK1 = sK2(sK0) <- (22) [avatar component clause 285] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
288. 22 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 283,169,164,159,154,149,285]
290. ~l1_collsp(sK0) | ~v4_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) | sK1 = sK3(sK0) [resolution 126,105] {a:1,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
292. ~v4_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) | sK1 = sK3(sK0) <- (2) [subsumption resolution 290,151] {a:1,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
293. ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) | sK1 = sK3(sK0) <- (2, 3) [subsumption resolution 292,156] {a:1,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
294. ~v2_collsp(sK0) | v2_struct_0(sK0) | sK1 = sK3(sK0) <- (2, 3, 4) [subsumption resolution 293,161] {a:1,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
295. v2_struct_0(sK0) | sK1 = sK3(sK0) <- (2, 3, 4, 5) [subsumption resolution 294,166] {a:1,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
296. sK1 = sK3(sK0) <- (2, 3, 4, 5, ~6) [subsumption resolution 295,171] {a:1,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:7,thDist:-7}
298. 23 <=> sK1 = sK3(sK0) [avatar definition]
300. sK1 = sK3(sK0) <- (23) [avatar component clause 298] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
301. 23 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 296,169,164,159,154,149,298]
315. sK1 != sK2(sK0) | ~l1_collsp(sK0) | ~v4_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (23) [superposition 128,300] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
317. ~l1_collsp(sK0) | ~v4_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (22, 23) [subsumption resolution 315,287] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
318. ~v4_collsp(sK0) | ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (2, 22, 23) [subsumption resolution 317,151] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
319. ~v3_collsp(sK0) | ~v2_collsp(sK0) | v2_struct_0(sK0) <- (2, 3, 22, 23) [subsumption resolution 318,156] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
320. ~v2_collsp(sK0) | v2_struct_0(sK0) <- (2, 3, 4, 22, 23) [subsumption resolution 319,161] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
321. v2_struct_0(sK0) <- (2, 3, 4, 5, 22, 23) [subsumption resolution 320,166] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
322. $false <- (2, 3, 4, 5, ~6, 22, 23) [subsumption resolution 321,171] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
323. ~2 | ~3 | ~4 | ~5 | 6 | ~22 | ~23 [avatar contradiction clause 322]
324. $false [avatar sat refutation 152,157,162,167,172,288,301,323] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:146,thDist:-146}
% SZS output end Proof for incproj__t7_incproj
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 34
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 44
% Generated clauses: 126
% Active clauses: 45
% Passive clauses: 60
% Final active clauses: 45
% Final passive clauses: 11
% Main loop iterations started: 49
% 
% >>> Simplifying Inferences
% Duplicate literals: 6
% Fw subsumption resolutions: 21
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Equational tautologies: 1
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 28
% Forward superposition: 2
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 25
% SAT solver unit clauses: 10
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

