i: [1,23,0,2,1,1,0,0] t38_rewrite1
a: [3,44,0,2,1,23]
s: 44
i: [1,24,0,8,2,1,0,0] t38_rewrite1
a: [3,35,0,4,1,24]
s: 35
i: [1,25,0,6,2,0,0,2] t12_rewrite1
s: 25
i: [1,26,0,13,3,0,0,1] d7_rewrite1
s: 26
i: [1,27,0,13,3,0,0,1] d7_rewrite1
s: 27
i: [1,29,0,13,3,0,0,1] d8_rewrite1
s: 29
i: [1,30,0,13,3,0,0,1] d8_rewrite1
s: 30
i: [1,28,0,14,4,0,0,1] d7_rewrite1
s: 28
d: [2,46,1,12,4,0,40,28,25]
d: [2,55,1,10,3,0,18,46]
s: 55
i: [1,31,0,14,4,0,0,1] d8_rewrite1
s: 31
d: [2,66,1,12,4,0,40,31,25]
d: [2,75,1,10,3,0,18,66]
s: 75
d: [2,56,2,8,3,0,40,55,25]
d: [2,65,2,6,2,0,18,56]
s: 65
d: [2,86,3,0,0,2,52,44,35,65]
e: 86
d: [2,87,3,2,1,1,40,65,35]
d: [2,89,3,0,0,2,21,87,44]
e: 89
r: 35
a: [3,34,0,4,1,24]
s: 34
a: [3,39,0,4,1,24]
s: 39
d: [2,50,1,19,5,0,40,28,30]
d: [2,51,1,17,4,0,18,50]
s: 51
d: [2,49,1,19,5,0,40,28,29]
d: [2,52,1,17,4,0,18,49]
s: 52
d: [2,76,2,8,3,0,40,75,25]
d: [2,85,2,6,2,0,18,76]
s: 85
d: [2,91,3,0,0,2,52,44,39,85]
e: 91
d: [2,92,3,2,1,1,40,85,39]
d: [2,94,3,0,0,2,21,92,44]
e: 94
f: 94,89,23,24
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rewrite1__t38_rewrite1
% SZS output start Proof for rewrite1__t38_rewrite1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r6_rewrite1(X0,X1,X1) & r5_rewrite1(X0,X1,X1))) [input t38_rewrite1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (r6_rewrite1(X0,X1,X1) & r5_rewrite1(X0,X1,X1))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r5_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3)))) [input d7_rewrite1]
4. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r6_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1)))) [input d8_rewrite1]
5. ! [X0] : (v1_relat_1(X0) => ! [X1] : r1_rewrite1(X0,X1,X1)) [input t12_rewrite1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r5_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3)))) [flattening 3]
7. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r6_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1)))) [flattening 4]
8. ? [X0] : (? [X1] : (~r6_rewrite1(X0,X1,X1) | ~r5_rewrite1(X0,X1,X1)) & v1_relat_1(X0)) [ennf transformation 2]
9. ! [X0] : (! [X1] : r1_rewrite1(X0,X1,X1) | ~v1_relat_1(X0)) [ennf transformation 5]
10. ! [X0] : (! [X1,X2] : (r5_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3))) | ~v1_relat_1(X0)) [ennf transformation 6]
11. ! [X0] : (! [X1,X2] : (r6_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1))) | ~v1_relat_1(X0)) [ennf transformation 7]
12. ? [X0] : (? [X1] : (~r6_rewrite1(X0,X1,X1) | ~r5_rewrite1(X0,X1,X1)) & v1_relat_1(X0)) => (? [X1] : (~r6_rewrite1(sK0,X1,X1) | ~r5_rewrite1(sK0,X1,X1)) & v1_relat_1(sK0)) [choice axiom]
13. ? [X1] : (~r6_rewrite1(sK0,X1,X1) | ~r5_rewrite1(sK0,X1,X1)) => (~r6_rewrite1(sK0,sK1,sK1) | ~r5_rewrite1(sK0,sK1,sK1)) [choice axiom]
14. (~r6_rewrite1(sK0,sK1,sK1) | ~r5_rewrite1(sK0,sK1,sK1)) & v1_relat_1(sK0) [skolemisation 8,13,12]
15. ! [X0] : (! [X1,X2] : ((r5_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X3))) & (? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3)) | ~r5_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [nnf transformation 10]
16. ! [X0] : (! [X1,X2] : ((r5_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X3))) & (? [X4] : (r1_rewrite1(X0,X2,X4) & r1_rewrite1(X0,X1,X4)) | ~r5_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [rectify 15]
17. ! [X2,X1,X0] : (? [X4] : (r1_rewrite1(X0,X2,X4) & r1_rewrite1(X0,X1,X4)) => (r1_rewrite1(X0,X2,sK2(X0,X1,X2)) & r1_rewrite1(X0,X1,sK2(X0,X1,X2)))) [choice axiom]
18. ! [X0] : (! [X1,X2] : ((r5_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X3))) & ((r1_rewrite1(X0,X2,sK2(X0,X1,X2)) & r1_rewrite1(X0,X1,sK2(X0,X1,X2))) | ~r5_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [skolemisation 16,17]
19. ! [X0] : (! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X3,X2) | ~r1_rewrite1(X0,X3,X1))) & (? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1)) | ~r6_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [nnf transformation 11]
20. ! [X0] : (! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X3,X2) | ~r1_rewrite1(X0,X3,X1))) & (? [X4] : (r1_rewrite1(X0,X4,X2) & r1_rewrite1(X0,X4,X1)) | ~r6_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [rectify 19]
21. ! [X2,X1,X0] : (? [X4] : (r1_rewrite1(X0,X4,X2) & r1_rewrite1(X0,X4,X1)) => (r1_rewrite1(X0,sK3(X0,X1,X2),X2) & r1_rewrite1(X0,sK3(X0,X1,X2),X1))) [choice axiom]
22. ! [X0] : (! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X3,X2) | ~r1_rewrite1(X0,X3,X1))) & ((r1_rewrite1(X0,sK3(X0,X1,X2),X2) & r1_rewrite1(X0,sK3(X0,X1,X2),X1)) | ~r6_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [skolemisation 20,21]
23. v1_relat_1(sK0) [cnf transformation 14] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
24. ~r6_rewrite1(sK0,sK1,sK1) | ~r5_rewrite1(sK0,sK1,sK1) [cnf transformation 14] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
25. r1_rewrite1(X0,X1,X1) | ~v1_relat_1(X0) [cnf transformation 9] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
28. ~r1_rewrite1(X0,X2,X3) | r5_rewrite1(X0,X1,X2) | ~r1_rewrite1(X0,X1,X3) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
31. ~r1_rewrite1(X0,X3,X2) | r6_rewrite1(X0,X1,X2) | ~r1_rewrite1(X0,X3,X1) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
33. 1 <=> r5_rewrite1(sK0,sK1,sK1) [avatar definition]
35. ~r5_rewrite1(sK0,sK1,sK1) <- (~1) [avatar component clause 33] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
37. 2 <=> r6_rewrite1(sK0,sK1,sK1) [avatar definition]
39. ~r6_rewrite1(sK0,sK1,sK1) <- (~2) [avatar component clause 37] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
40. ~1 | ~2 [avatar split clause 24,37,33]
42. 3 <=> v1_relat_1(sK0) [avatar definition]
44. v1_relat_1(sK0) <- (3) [avatar component clause 42] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
45. 3 [avatar split clause 23,42]
46. r5_rewrite1(X0,X1,X2) | ~r1_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) | ~v1_relat_1(X0) [resolution 28,25] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:3,thDist:-3}
55. ~r1_rewrite1(X0,X1,X2) | r5_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [duplicate literal removal 46] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
56. r5_rewrite1(X0,X1,X1) | ~v1_relat_1(X0) | ~v1_relat_1(X0) [resolution 55,25] {a:2,w:8,wCS:3200,sine:1,thAx:0,allAx:4,thDist:-4}
65. r5_rewrite1(X0,X1,X1) | ~v1_relat_1(X0) [duplicate literal removal 56] {a:2,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
66. r6_rewrite1(X0,X1,X2) | ~r1_rewrite1(X0,X2,X1) | ~v1_relat_1(X0) | ~v1_relat_1(X0) [resolution 31,25] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:3,thDist:-3}
75. ~r1_rewrite1(X0,X2,X1) | r6_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [duplicate literal removal 66] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
76. r6_rewrite1(X0,X1,X1) | ~v1_relat_1(X0) | ~v1_relat_1(X0) [resolution 75,25] {a:2,w:8,wCS:3200,sine:1,thAx:0,allAx:4,thDist:-4}
85. r6_rewrite1(X0,X1,X1) | ~v1_relat_1(X0) [duplicate literal removal 76] {a:2,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
87. ~v1_relat_1(sK0) <- (~1) [resolution 65,35] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
89. $false <- (~1, 3) [subsumption resolution 87,44] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
90. 1 | ~3 [avatar contradiction clause 89]
92. ~v1_relat_1(sK0) <- (~2) [resolution 85,39] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
94. $false <- (~2, 3) [subsumption resolution 92,44] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
95. 2 | ~3 [avatar contradiction clause 94]
96. $false [avatar sat refutation 40,45,90,95] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:43,thDist:-43}
% SZS output end Proof for rewrite1__t38_rewrite1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 4
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 9
% Generated clauses: 59
% Active clauses: 17
% Passive clauses: 31
% Final active clauses: 16
% Final passive clauses: 14
% Main loop iterations started: 17
% 
% >>> Simplifying Inferences
% Duplicate literals: 20
% Fw subsumption resolutions: 2
% 
% >>> Generating Inferences
% Binary resolution: 22
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 8
% SAT solver unit clauses: 1
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

