i: [1,23,0,2,1,1,0,0] t157_relat_1
a: [3,40,0,2,1,23]
s: 40
i: [1,24,0,10,1,1,0,0] t157_relat_1
a: [3,35,0,10,1,24]
s: 35
i: [1,27,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 27
i: [1,29,0,6,2,0,0,1] fc1_relat_1
s: 29
i: [1,30,0,6,2,0,0,1] dt_k5_relat_1
s: 30
i: [1,28,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 28
r: 35
d: [2,43,0,10,1,1,23,35,28]
a: [3,49,0,10,1,43]
s: 49
i: [1,25,0,8,2,0,0,1] rd4_relat_1
s: 25
d: [2,44,1,6,2,0,44,29,28]
s: 44
i: [1,31,0,11,2,0,0,1] rd5_relat_1
s: 31
i: [1,26,0,13,2,0,0,1] t71_relat_1
s: 26
d: [2,45,1,6,2,0,44,29,28]
s: 45
d: [2,60,1,14,3,0,44,26,25]
d: [2,71,1,12,2,0,18,60]
s: 71
d: [2,77,2,12,2,0,44,71,28]
s: 77
d: [2,89,3,0,0,2,52,40,49,77]
e: 89
d: [2,96,3,9,2,1,44,49,77]
d: [2,98,3,2,1,1,19,96]
d: [2,104,3,0,0,2,21,98,40]
e: 104
f: 104,43,23,24
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t157_relat_1
% SZS output start Proof for relat_1__t157_relat_1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => k5_relat_1(X1,X0) = k5_relat_1(X1,k3_xboole_0(k9_xtuple_0(X1),X0))) [input t157_relat_1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => k5_relat_1(X1,X0) = k5_relat_1(X1,k3_xboole_0(k9_xtuple_0(X1),X0))) [negated conjecture 1]
3. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
9. ! [X0] : (v1_relat_1(X0) => k5_relat_1(X0,k9_xtuple_0(X0)) = X0) [input rd4_relat_1]
11. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1))) [input t71_relat_1]
12. ~! [X0,X1] : (v1_relat_1(X1) => k5_relat_1(X1,X0) = k5_relat_1(X1,k3_xboole_0(k9_xtuple_0(X1),X0))) [flattening 2]
13. ! [X0,X1,X2] : (v1_relat_1(X2) => k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1))) [flattening 11]
15. ? [X0,X1] : (k5_relat_1(X1,X0) != k5_relat_1(X1,k3_xboole_0(k9_xtuple_0(X1),X0)) & v1_relat_1(X1)) [ennf transformation 12]
16. ! [X0] : (k5_relat_1(X0,k9_xtuple_0(X0)) = X0 | ~v1_relat_1(X0)) [ennf transformation 9]
17. ! [X0,X1,X2] : (k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1)) | ~v1_relat_1(X2)) [ennf transformation 13]
21. ? [X0,X1] : (k5_relat_1(X1,X0) != k5_relat_1(X1,k3_xboole_0(k9_xtuple_0(X1),X0)) & v1_relat_1(X1)) => (k5_relat_1(sK1,sK0) != k5_relat_1(sK1,k3_xboole_0(k9_xtuple_0(sK1),sK0)) & v1_relat_1(sK1)) [choice axiom]
22. k5_relat_1(sK1,sK0) != k5_relat_1(sK1,k3_xboole_0(k9_xtuple_0(sK1),sK0)) & v1_relat_1(sK1) [skolemisation 15,21]
23. v1_relat_1(sK1) [cnf transformation 22] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
24. k5_relat_1(sK1,sK0) != k5_relat_1(sK1,k3_xboole_0(k9_xtuple_0(sK1),sK0)) [cnf transformation 22] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
25. k5_relat_1(X0,k9_xtuple_0(X0)) = X0 | ~v1_relat_1(X0) [cnf transformation 16] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
26. k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1)) | ~v1_relat_1(X2) [cnf transformation 17] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
28. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. 1 <=> k5_relat_1(sK1,sK0) = k5_relat_1(sK1,k3_xboole_0(k9_xtuple_0(sK1),sK0)) [avatar definition]
35. k5_relat_1(sK1,sK0) != k5_relat_1(sK1,k3_xboole_0(k9_xtuple_0(sK1),sK0)) <- (~1) [avatar component clause 33] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
36. ~1 [avatar split clause 24,33]
38. 2 <=> v1_relat_1(sK1) [avatar definition]
40. v1_relat_1(sK1) <- (2) [avatar component clause 38] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
41. 2 [avatar split clause 23,38]
43. k5_relat_1(sK1,sK0) != k5_relat_1(sK1,k3_xboole_0(sK0,k9_xtuple_0(sK1))) <- (~1) [backward demodulation 35,28] {a:0,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
47. 3 <=> k5_relat_1(sK1,sK0) = k5_relat_1(sK1,k3_xboole_0(sK0,k9_xtuple_0(sK1))) [avatar definition]
49. k5_relat_1(sK1,sK0) != k5_relat_1(sK1,k3_xboole_0(sK0,k9_xtuple_0(sK1))) <- (~3) [avatar component clause 47] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
50. ~3 | 1 [avatar split clause 43,33,47]
60. k5_relat_1(X0,X1) = k5_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),X1)) | ~v1_relat_1(X0) | ~v1_relat_1(X0) [superposition 26,25] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
71. k5_relat_1(X0,X1) = k5_relat_1(X0,k3_xboole_0(k9_xtuple_0(X0),X1)) | ~v1_relat_1(X0) [duplicate literal removal 60] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
77. k5_relat_1(X1,X2) = k5_relat_1(X1,k3_xboole_0(X2,k9_xtuple_0(X1))) | ~v1_relat_1(X1) [superposition 71,28] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
96. k5_relat_1(sK1,sK0) != k5_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~3) [superposition 49,77] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
98. ~v1_relat_1(sK1) <- (~3) [trivial inequality removal 96] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
104. $false <- (2, ~3) [subsumption resolution 98,40] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
105. ~2 | 3 [avatar contradiction clause 104]
107. $false [avatar sat refutation 36,41,50,105] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:22,thDist:-22}
% SZS output end Proof for relat_1__t157_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 9
% Generated clauses: 70
% Active clauses: 13
% Passive clauses: 24
% Final active clauses: 12
% Final passive clauses: 10
% Main loop iterations started: 14
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Trivial inequalities: 1
% Fw subsumption resolutions: 10
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 10
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Unit resulting resolution: 1
% Forward superposition: 15
% Backward superposition: 17
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 7
% SAT solver unit clauses: 2
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

