i: [1,19,0,2,1,1,0,0] t43_rewrite1
a: [3,57,0,2,1,19]
s: 57
i: [1,24,0,4,1,1,0,0] t43_rewrite1
a: [3,32,0,4,1,24]
s: 32
i: [1,23,0,8,2,1,0,0] t43_rewrite1
a: [3,41,0,4,1,23]
s: 41
i: [1,21,0,8,2,1,0,0] t43_rewrite1
a: [3,51,0,4,1,21]
s: 51
i: [1,25,0,13,3,0,0,1] d8_rewrite1
s: 25
i: [1,26,0,13,3,0,0,1] d8_rewrite1
s: 26
i: [1,27,0,14,4,0,0,1] d8_rewrite1
s: 27
i: [1,28,0,14,4,0,0,1] t16_rewrite1
s: 28
d: [2,59,1,10,3,1,40,27,41]
d: [2,64,1,8,2,2,21,59,57]
s: 64
d: [2,65,1,10,3,1,40,28,41]
d: [2,70,1,8,2,2,21,65,57]
s: 70
d: [2,61,1,19,5,0,40,27,26]
d: [2,62,1,17,4,0,18,61]
s: 62
d: [2,71,2,4,1,2,40,64,41]
a: [3,75,2,4,1,71]
s: 75
d: [2,60,1,19,5,0,40,27,25]
d: [2,63,1,17,4,0,18,60]
s: 63
d: [2,77,2,13,3,2,40,70,25]
d: [2,79,2,11,2,2,21,77,57]
s: 79
d: [2,67,1,19,5,0,40,28,26]
d: [2,68,1,17,4,0,18,67]
s: 68
d: [2,66,1,19,5,0,40,28,25]
d: [2,69,1,17,4,0,18,66]
s: 69
d: [2,89,3,14,4,2,40,79,63]
d: [2,94,3,10,3,2,18,89]
d: [2,95,3,8,2,2,21,94,57]
a: [3,99,3,4,1,95]
s: 99
d: [2,78,2,13,3,2,40,70,26]
d: [2,80,2,11,2,2,21,78,57]
s: 80
d: [2,90,3,14,4,2,40,79,62]
d: [2,93,3,10,3,2,18,90]
d: [2,104,3,8,2,2,21,93,57]
s: 104
d: [2,82,2,16,5,0,40,62,26]
d: [2,83,2,10,3,0,18,82]
s: 83
d: [2,86,2,16,5,0,40,63,26]
d: [2,87,2,10,3,0,18,86]
s: 87
r: 51
a: [3,37,0,4,1,23]
s: 37
i: [1,22,0,8,2,1,0,0] t43_rewrite1
a: [3,46,0,4,1,22]
s: 46
d: [2,149,2,0,0,4,52,32,46,64]
e: 149
d: [2,150,1,0,0,4,52,57,41,32,46,27]
e: 150
d: [2,151,2,4,1,3,40,46,64]
d: [2,156,2,0,0,4,21,151,32]
e: 156
r: 104
r: 80
r: 79
r: 70
r: 64
r: 41
a: [3,50,0,4,1,21]
s: 50
d: [2,152,1,10,3,1,40,46,28]
d: [2,158,1,8,2,2,21,152,57]
s: 158
d: [2,153,1,10,3,1,40,46,27]
d: [2,159,1,8,2,2,21,153,57]
s: 159
r: 159
r: 158
r: 46
r: 50
s: 41
a: [3,45,0,4,1,22]
s: 45
s: 51
d: [2,170,1,10,3,1,40,41,28]
d: [2,172,1,8,2,2,21,170,57]
s: 172
d: [2,171,1,10,3,1,40,41,27]
d: [2,173,1,8,2,2,21,171,57]
s: 173
d: [2,124,3,6,2,1,40,83,51]
d: [2,127,3,4,1,2,21,124,57]
a: [3,131,3,4,1,127]
s: 131
d: [2,180,2,13,3,2,40,172,25]
d: [2,182,2,11,2,2,21,180,57]
s: 182
d: [2,135,3,6,2,1,40,87,51]
d: [2,138,3,4,1,2,21,135,57]
a: [3,142,3,4,1,138]
s: 142
d: [2,164,2,4,1,2,40,159,46]
a: [3,168,2,4,1,164]
s: 168
d: [2,181,2,13,3,2,40,172,26]
d: [2,183,2,11,2,2,21,181,57]
s: 183
d: [2,188,3,14,4,2,40,182,62]
d: [2,191,3,10,3,2,18,188]
d: [2,193,3,8,2,2,21,191,57]
s: 193
d: [2,85,2,16,5,0,40,63,25]
d: [2,88,2,10,3,0,18,85]
s: 88
d: [2,213,4,4,1,3,40,193,142]
a: [3,217,4,4,1,213]
s: 217
d: [2,226,5,0,0,3,52,57,32,217,87]
e: 226
d: [2,228,5,6,2,1,40,217,87]
d: [2,231,5,4,1,2,21,228,57]
d: [2,232,5,0,0,3,21,231,32]
e: 232
r: 193
r: 183
r: 182
r: 173
r: 172
r: 41
r: 45
r: 217
a: [3,40,0,4,1,23]
s: 40
s: 46
a: [3,216,4,4,1,213]
s: 216
d: [2,234,1,10,3,1,40,46,28]
d: [2,236,1,8,2,2,21,234,57]
s: 236
d: [2,235,1,10,3,1,40,46,27]
d: [2,237,1,8,2,2,21,235,57]
s: 237
d: [2,238,2,13,3,2,40,236,25]
d: [2,240,2,11,2,2,21,238,57]
s: 240
d: [2,239,2,13,3,2,40,236,26]
d: [2,241,2,11,2,2,21,239,57]
s: 241
d: [2,244,3,14,4,2,40,240,62]
d: [2,247,3,10,3,2,18,244]
d: [2,249,3,8,2,2,21,247,57]
s: 249
d: [2,261,4,0,0,4,52,32,37,249]
e: 261
d: [2,262,4,4,1,3,40,249,37]
d: [2,266,4,0,0,4,21,262,32]
e: 266
i: [1,20,0,8,2,1,0,0] t43_rewrite1
d: [2,174,3,6,2,1,40,51,87]
d: [2,176,3,4,1,2,21,174,57]
f: 266,232,213,176,19,20,21,22,23,24
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rewrite1__t43_rewrite1
% SZS output start Proof for rewrite1__t43_rewrite1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : (((r1_rewrite1(X0,X2,X3) & r6_rewrite1(X0,X1,X2)) | (r6_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X2,X1))) => r6_rewrite1(X0,X1,X3))) [input t43_rewrite1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : (((r1_rewrite1(X0,X2,X3) & r6_rewrite1(X0,X1,X2)) | (r6_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X2,X1))) => r6_rewrite1(X0,X1,X3))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r6_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1)))) [input d8_rewrite1]
4. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : ((r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2)) => r1_rewrite1(X0,X1,X3))) [input t16_rewrite1]
5. ~! [X0] : (v1_relat_1(X0) => ! [X1,X2,X3] : (((r1_rewrite1(X0,X2,X3) & r6_rewrite1(X0,X1,X2)) | (r6_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X2,X1))) => r6_rewrite1(X0,X1,X3))) [flattening 2]
6. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r6_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1)))) [flattening 3]
7. ! [X0] : (v1_relat_1(X0) => ! [X1,X2,X3] : ((r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X2)) => r1_rewrite1(X0,X1,X3))) [flattening 4]
8. ? [X0] : (? [X1,X2,X3] : (~r6_rewrite1(X0,X1,X3) & ((r1_rewrite1(X0,X2,X3) & r6_rewrite1(X0,X1,X2)) | (r6_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X2,X1)))) & v1_relat_1(X0)) [ennf transformation 5]
9. ! [X0] : (! [X1,X2] : (r6_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1))) | ~v1_relat_1(X0)) [ennf transformation 6]
10. ! [X0] : (! [X1,X2,X3] : (r1_rewrite1(X0,X1,X3) | (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [ennf transformation 7]
11. ! [X0] : (! [X1,X2,X3] : (r1_rewrite1(X0,X1,X3) | ~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X2)) | ~v1_relat_1(X0)) [flattening 10]
12. ? [X0] : (? [X1,X2,X3] : (~r6_rewrite1(X0,X1,X3) & ((r1_rewrite1(X0,X2,X3) & r6_rewrite1(X0,X1,X2)) | (r6_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X2,X1)))) & v1_relat_1(X0)) => (? [X3,X2,X1] : (~r6_rewrite1(sK0,X1,X3) & ((r1_rewrite1(sK0,X2,X3) & r6_rewrite1(sK0,X1,X2)) | (r6_rewrite1(sK0,X2,X3) & r1_rewrite1(sK0,X2,X1)))) & v1_relat_1(sK0)) [choice axiom]
13. ? [X3,X2,X1] : (~r6_rewrite1(sK0,X1,X3) & ((r1_rewrite1(sK0,X2,X3) & r6_rewrite1(sK0,X1,X2)) | (r6_rewrite1(sK0,X2,X3) & r1_rewrite1(sK0,X2,X1)))) => (~r6_rewrite1(sK0,sK1,sK3) & ((r1_rewrite1(sK0,sK2,sK3) & r6_rewrite1(sK0,sK1,sK2)) | (r6_rewrite1(sK0,sK2,sK3) & r1_rewrite1(sK0,sK2,sK1)))) [choice axiom]
14. (~r6_rewrite1(sK0,sK1,sK3) & ((r1_rewrite1(sK0,sK2,sK3) & r6_rewrite1(sK0,sK1,sK2)) | (r6_rewrite1(sK0,sK2,sK3) & r1_rewrite1(sK0,sK2,sK1)))) & v1_relat_1(sK0) [skolemisation 8,13,12]
15. ! [X0] : (! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X3,X2) | ~r1_rewrite1(X0,X3,X1))) & (? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1)) | ~r6_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [nnf transformation 9]
16. ! [X0] : (! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X3,X2) | ~r1_rewrite1(X0,X3,X1))) & (? [X4] : (r1_rewrite1(X0,X4,X2) & r1_rewrite1(X0,X4,X1)) | ~r6_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [rectify 15]
17. ! [X2,X1,X0] : (? [X4] : (r1_rewrite1(X0,X4,X2) & r1_rewrite1(X0,X4,X1)) => (r1_rewrite1(X0,sK4(X0,X1,X2),X2) & r1_rewrite1(X0,sK4(X0,X1,X2),X1))) [choice axiom]
18. ! [X0] : (! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X3,X2) | ~r1_rewrite1(X0,X3,X1))) & ((r1_rewrite1(X0,sK4(X0,X1,X2),X2) & r1_rewrite1(X0,sK4(X0,X1,X2),X1)) | ~r6_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [skolemisation 16,17]
19. v1_relat_1(sK0) [cnf transformation 14] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
20. r6_rewrite1(sK0,sK1,sK2) | r1_rewrite1(sK0,sK2,sK1) [cnf transformation 14] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
21. r6_rewrite1(sK0,sK1,sK2) | r6_rewrite1(sK0,sK2,sK3) [cnf transformation 14] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
22. r1_rewrite1(sK0,sK2,sK3) | r1_rewrite1(sK0,sK2,sK1) [cnf transformation 14] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
23. r1_rewrite1(sK0,sK2,sK3) | r6_rewrite1(sK0,sK2,sK3) [cnf transformation 14] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
24. ~r6_rewrite1(sK0,sK1,sK3) [cnf transformation 14] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
25. r1_rewrite1(X0,sK4(X0,X1,X2),X1) | ~r6_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
26. r1_rewrite1(X0,sK4(X0,X1,X2),X2) | ~r6_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
27. ~r1_rewrite1(X0,X3,X2) | r6_rewrite1(X0,X1,X2) | ~r1_rewrite1(X0,X3,X1) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
28. ~r1_rewrite1(X0,X2,X3) | r1_rewrite1(X0,X1,X3) | ~r1_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 11] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
30. 1 <=> r6_rewrite1(sK0,sK1,sK3) [avatar definition]
32. ~r6_rewrite1(sK0,sK1,sK3) <- (~1) [avatar component clause 30] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
33. ~1 [avatar split clause 24,30]
35. 2 <=> r6_rewrite1(sK0,sK2,sK3) [avatar definition]
37. r6_rewrite1(sK0,sK2,sK3) <- (2) [avatar component clause 35] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
39. 3 <=> r1_rewrite1(sK0,sK2,sK3) [avatar definition]
41. r1_rewrite1(sK0,sK2,sK3) <- (3) [avatar component clause 39] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
42. 2 | 3 [avatar split clause 23,39,35]
44. 4 <=> r1_rewrite1(sK0,sK2,sK1) [avatar definition]
46. r1_rewrite1(sK0,sK2,sK1) <- (4) [avatar component clause 44] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
47. 4 | 3 [avatar split clause 22,39,44]
49. 5 <=> r6_rewrite1(sK0,sK1,sK2) [avatar definition]
51. r6_rewrite1(sK0,sK1,sK2) <- (5) [avatar component clause 49] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
52. 2 | 5 [avatar split clause 21,49,35]
53. 4 | 5 [avatar split clause 20,49,44]
55. 6 <=> v1_relat_1(sK0) [avatar definition]
57. v1_relat_1(sK0) <- (6) [avatar component clause 55] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
58. 6 [avatar split clause 19,55]
60. r6_rewrite1(X1,X2,X3) | ~r1_rewrite1(X1,sK4(X1,X3,X4),X2) | ~v1_relat_1(X1) | ~r6_rewrite1(X1,X3,X4) | ~v1_relat_1(X1) [resolution 27,25] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
61. r6_rewrite1(X5,X6,X7) | ~r1_rewrite1(X5,sK4(X5,X8,X7),X6) | ~v1_relat_1(X5) | ~r6_rewrite1(X5,X8,X7) | ~v1_relat_1(X5) [resolution 27,26] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
62. ~r1_rewrite1(X5,sK4(X5,X8,X7),X6) | r6_rewrite1(X5,X6,X7) | ~v1_relat_1(X5) | ~r6_rewrite1(X5,X8,X7) [duplicate literal removal 61] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
63. ~r1_rewrite1(X1,sK4(X1,X3,X4),X2) | r6_rewrite1(X1,X2,X3) | ~v1_relat_1(X1) | ~r6_rewrite1(X1,X3,X4) [duplicate literal removal 60] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
86. r6_rewrite1(X3,X4,X5) | ~v1_relat_1(X3) | ~r6_rewrite1(X3,X5,X4) | ~r6_rewrite1(X3,X5,X4) | ~v1_relat_1(X3) [resolution 63,26] {a:2,w:16,wCS:6400,sine:1,thAx:0,allAx:6,thDist:-6}
87. ~r6_rewrite1(X3,X5,X4) | ~v1_relat_1(X3) | r6_rewrite1(X3,X4,X5) [duplicate literal removal 86] {a:2,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
140. 11 <=> r6_rewrite1(sK0,sK2,sK1) [avatar definition]
142. r6_rewrite1(sK0,sK2,sK1) <- (11) [avatar component clause 140] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
170. r1_rewrite1(sK0,X0,sK3) | ~r1_rewrite1(sK0,X0,sK2) | ~v1_relat_1(sK0) <- (3) [resolution 41,28] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
172. ~r1_rewrite1(sK0,X0,sK2) | r1_rewrite1(sK0,X0,sK3) <- (3, 6) [subsumption resolution 170,57] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
174. ~v1_relat_1(sK0) | r6_rewrite1(sK0,sK2,sK1) <- (5) [resolution 51,87] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
176. r6_rewrite1(sK0,sK2,sK1) <- (5, 6) [subsumption resolution 174,57] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
177. 11 | ~5 | ~6 [avatar split clause 176,55,49,140]
180. r1_rewrite1(sK0,sK4(sK0,sK2,X0),sK3) | ~r6_rewrite1(sK0,sK2,X0) | ~v1_relat_1(sK0) <- (3, 6) [resolution 172,25] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
182. r1_rewrite1(sK0,sK4(sK0,sK2,X0),sK3) | ~r6_rewrite1(sK0,sK2,X0) <- (3, 6) [subsumption resolution 180,57] {a:2,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
188. ~r6_rewrite1(sK0,sK2,X1) | r6_rewrite1(sK0,sK3,X1) | ~v1_relat_1(sK0) | ~r6_rewrite1(sK0,sK2,X1) <- (3, 6) [resolution 182,62] {a:3,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
191. ~r6_rewrite1(sK0,sK2,X1) | r6_rewrite1(sK0,sK3,X1) | ~v1_relat_1(sK0) <- (3, 6) [duplicate literal removal 188] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
193. ~r6_rewrite1(sK0,sK2,X1) | r6_rewrite1(sK0,sK3,X1) <- (3, 6) [subsumption resolution 191,57] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
213. r6_rewrite1(sK0,sK3,sK1) <- (3, 6, 11) [resolution 193,142] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:19,thDist:-19}
215. 13 <=> r6_rewrite1(sK0,sK3,sK1) [avatar definition]
217. r6_rewrite1(sK0,sK3,sK1) <- (13) [avatar component clause 215] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:19,thDist:-19}
218. 13 | ~3 | ~6 | ~11 [avatar split clause 213,140,55,39,215]
228. ~v1_relat_1(sK0) | r6_rewrite1(sK0,sK1,sK3) <- (13) [resolution 217,87] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:25,thDist:-25}
231. r6_rewrite1(sK0,sK1,sK3) <- (6, 13) [subsumption resolution 228,57] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:25,thDist:-25}
232. $false <- (~1, 6, 13) [subsumption resolution 231,32] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:25,thDist:-25}
233. 1 | ~6 | ~13 [avatar contradiction clause 232]
234. r1_rewrite1(sK0,X0,sK1) | ~r1_rewrite1(sK0,X0,sK2) | ~v1_relat_1(sK0) <- (4) [resolution 46,28] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
236. ~r1_rewrite1(sK0,X0,sK2) | r1_rewrite1(sK0,X0,sK1) <- (4, 6) [subsumption resolution 234,57] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
238. r1_rewrite1(sK0,sK4(sK0,sK2,X0),sK1) | ~r6_rewrite1(sK0,sK2,X0) | ~v1_relat_1(sK0) <- (4, 6) [resolution 236,25] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
240. r1_rewrite1(sK0,sK4(sK0,sK2,X0),sK1) | ~r6_rewrite1(sK0,sK2,X0) <- (4, 6) [subsumption resolution 238,57] {a:2,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
244. ~r6_rewrite1(sK0,sK2,X1) | r6_rewrite1(sK0,sK1,X1) | ~v1_relat_1(sK0) | ~r6_rewrite1(sK0,sK2,X1) <- (4, 6) [resolution 240,62] {a:3,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
247. ~r6_rewrite1(sK0,sK2,X1) | r6_rewrite1(sK0,sK1,X1) | ~v1_relat_1(sK0) <- (4, 6) [duplicate literal removal 244] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
249. ~r6_rewrite1(sK0,sK2,X1) | r6_rewrite1(sK0,sK1,X1) <- (4, 6) [subsumption resolution 247,57] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
262. r6_rewrite1(sK0,sK1,sK3) <- (2, 4, 6) [resolution 249,37] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
266. $false <- (~1, 2, 4, 6) [subsumption resolution 262,32] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
267. 1 | ~2 | ~4 | ~6 [avatar contradiction clause 266]
268. $false [avatar sat refutation 33,42,47,52,53,58,177,218,233,267] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:309,thDist:-309}
% SZS output end Proof for rewrite1__t43_rewrite1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 3
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 10
% Generated clauses: 195
% Active clauses: 47
% Passive clauses: 68
% Final active clauses: 28
% Final passive clauses: 9
% Main loop iterations started: 47
% 
% >>> Simplifying Inferences
% Duplicate literals: 28
% Fw subsumption resolutions: 46
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 30
% 
% >>> Generating Inferences
% Binary resolution: 90
% Unit resulting resolution: 4
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 2
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% resolution: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

