i: [1,48,0,9,3,1,0,0] t3_xtuple_0
a: [3,66,0,3,1,48]
s: 66
i: [1,49,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,84,0,2,1,49]
s: 84
i: [1,50,0,2,1,0,0,255] fc1_xboole_0
a: [3,89,0,2,1,50]
s: 89
i: [1,61,0,2,1,0,0,2] rc2_xboole_0
a: [3,99,0,2,1,61]
s: 99
i: [1,62,0,2,1,0,0,2] rc1_xboole_0
a: [3,104,0,2,1,62]
s: 104
i: [1,47,0,9,1,1,0,0] t3_xtuple_0
a: [3,79,0,9,1,47]
s: 79
i: [1,52,0,3,1,0,0,255] fc2_xboole_0
s: 52
i: [1,51,0,3,1,0,0,255] d2_xboole_0
a: [3,94,0,3,1,51]
s: 94
r: 84
i: [1,59,0,4,1,0,0,255] fc3_xboole_0
s: 59
i: [1,53,0,5,2,0,0,255] t6_boole
s: 53
i: [1,55,0,7,3,0,0,1] t8_boole
s: 55
i: [1,60,0,7,1,0,0,255] commutativity_k2_tarski
s: 60
i: [1,54,0,10,1,0,0,255] d5_tarski
s: 54
d: [2,118,1,4,1,0,44,59,54]
s: 118
i: [1,56,0,10,1,0,0,1] d4_xtuple_0
s: 56
i: [1,57,0,10,2,0,0,2] t1_xtuple_0
s: 57
d: [2,120,2,5,1,0,44,118,56]
s: 120
i: [1,58,0,10,2,0,0,2] t1_xtuple_0
s: 58
d: [2,107,1,7,3,0,44,53,53]
s: 107
d: [2,111,1,10,1,0,44,54,60]
s: 111
d: [2,112,1,10,1,0,44,54,60]
s: 112
d: [2,114,1,10,1,0,44,54,60]
s: 114
d: [2,115,1,10,1,0,44,54,60]
s: 115
d: [2,116,1,10,1,0,44,60,54]
s: 116
d: [2,117,1,10,1,0,44,60,54]
s: 117
d: [2,125,1,11,2,0,44,58,56]
s: 125
d: [2,146,2,11,2,1,44,125,79]
s: 146
d: [2,126,1,11,2,0,44,58,56]
s: 126
d: [2,148,3,12,2,1,44,146,56]
s: 148
a: [3,73,0,3,1,48]
s: 73
r: 148
r: 146
r: 79
d: [2,154,0,9,1,2,23,79,73]
a: [3,158,0,9,1,154]
s: 158
d: [2,152,3,12,2,2,23,148,73]
s: 152
d: [2,119,1,13,1,0,44,56,56]
s: 119
d: [2,132,2,10,1,0,44,111,60]
s: 132
d: [2,121,1,13,2,0,44,57,56]
s: 121
d: [2,122,1,13,2,0,44,57,56]
s: 122
d: [2,179,2,13,2,1,44,121,158]
s: 179
d: [2,113,1,15,1,0,44,54,54]
s: 113
d: [2,183,3,16,2,1,44,179,56]
s: 183
d: [2,133,2,10,1,0,44,111,60]
s: 133
d: [2,135,2,10,1,0,44,60,111]
s: 135
d: [2,207,4,7,1,1,47,183]
a: [3,212,4,7,1,207]
s: 212
r: 183
r: 179
d: [2,216,5,0,0,2,52,66,212,57]
e: 216
d: [2,217,5,0,0,2,52,66,212,57]
e: 217
r: 66
a: [3,65,0,3,1,48]
s: 65
r: 212
r: 158
a: [3,70,0,3,1,48]
s: 70
d: [2,236,0,9,1,2,23,158,65]
a: [3,245,0,9,1,236]
s: 245
d: [2,247,2,13,2,1,44,121,245]
s: 247
d: [2,235,4,7,1,2,23,212,65]
a: [3,240,4,7,1,235]
s: 240
r: 247
d: [2,253,5,0,0,2,52,70,240,58]
e: 253
d: [2,254,5,0,0,2,52,70,240,58]
e: 254
d: [2,150,4,3,1,1,47,148]
f: 254,235,217,207,154,150,47,48
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t3_xtuple_0
% SZS output start Proof for xtuple_0__t3_xtuple_0
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (k3_xtuple_0(X0,X1,X2) = k3_xtuple_0(X3,X4,X5) => (X2 = X5 & X1 = X4 & X0 = X3)) [input t3_xtuple_0]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (k3_xtuple_0(X0,X1,X2) = k3_xtuple_0(X3,X4,X5) => (X2 = X5 & X1 = X4 & X0 = X3)) [negated conjecture 1]
6. ! [X0] : ! [X1] : ! [X2] : k3_xtuple_0(X0,X1,X2) = k4_tarski(k4_tarski(X0,X1),X2) [input d4_xtuple_0]
22. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (k4_tarski(X0,X1) = k4_tarski(X2,X3) => (X1 = X3 & X0 = X2)) [input t1_xtuple_0]
26. ~! [X0,X1,X2,X3,X4,X5] : (k3_xtuple_0(X0,X1,X2) = k3_xtuple_0(X3,X4,X5) => (X2 = X5 & X1 = X4 & X0 = X3)) [flattening 2]
30. ! [X0,X1,X2] : k3_xtuple_0(X0,X1,X2) = k4_tarski(k4_tarski(X0,X1),X2) [flattening 6]
31. ! [X0,X1,X2,X3] : (k4_tarski(X0,X1) = k4_tarski(X2,X3) => (X1 = X3 & X0 = X2)) [flattening 22]
37. ? [X0,X1,X2,X3,X4,X5] : ((X2 != X5 | X1 != X4 | X0 != X3) & k3_xtuple_0(X0,X1,X2) = k3_xtuple_0(X3,X4,X5)) [ennf transformation 26]
40. ! [X0,X1,X2,X3] : ((X1 = X3 & X0 = X2) | k4_tarski(X0,X1) != k4_tarski(X2,X3)) [ennf transformation 31]
41. ? [X0,X1,X2,X3,X4,X5] : ((X2 != X5 | X1 != X4 | X0 != X3) & k3_xtuple_0(X0,X1,X2) = k3_xtuple_0(X3,X4,X5)) => ((sK2 != sK5 | sK1 != sK4 | sK0 != sK3) & k3_xtuple_0(sK0,sK1,sK2) = k3_xtuple_0(sK3,sK4,sK5)) [choice axiom]
42. (sK2 != sK5 | sK1 != sK4 | sK0 != sK3) & k3_xtuple_0(sK0,sK1,sK2) = k3_xtuple_0(sK3,sK4,sK5) [skolemisation 37,41]
47. k3_xtuple_0(sK0,sK1,sK2) = k3_xtuple_0(sK3,sK4,sK5) [cnf transformation 42] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. sK2 != sK5 | sK1 != sK4 | sK0 != sK3 [cnf transformation 42] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
56. k3_xtuple_0(X0,X1,X2) = k4_tarski(k4_tarski(X0,X1),X2) [cnf transformation 30] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. k4_tarski(X0,X1) != k4_tarski(X2,X3) | X0 = X2 [cnf transformation 40] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
58. k4_tarski(X0,X1) != k4_tarski(X2,X3) | X1 = X3 [cnf transformation 40] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
64. 1 <=> sK0 = sK3 [avatar definition]
65. sK0 = sK3 <- (1) [avatar component clause 64] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66. sK0 != sK3 <- (~1) [avatar component clause 64] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. 2 <=> sK1 = sK4 [avatar definition]
70. sK1 != sK4 <- (~2) [avatar component clause 68] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. 3 <=> sK2 = sK5 [avatar definition]
73. sK2 = sK5 <- (3) [avatar component clause 72] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. ~1 | ~2 | ~3 [avatar split clause 48,72,68,64]
77. 4 <=> k3_xtuple_0(sK0,sK1,sK2) = k3_xtuple_0(sK3,sK4,sK5) [avatar definition]
79. k3_xtuple_0(sK0,sK1,sK2) = k3_xtuple_0(sK3,sK4,sK5) <- (4) [avatar component clause 77] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. 4 [avatar split clause 47,77]
121. k3_xtuple_0(X0,X1,X2) != k4_tarski(X3,X4) | k4_tarski(X0,X1) = X3 [superposition 57,56] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
125. k3_xtuple_0(X0,X1,X2) != k4_tarski(X3,X4) | X2 = X4 [superposition 58,56] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
146. k4_tarski(X0,X1) != k3_xtuple_0(sK0,sK1,sK2) | sK5 = X1 <- (4) [superposition 125,79] {a:2,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
148. k3_xtuple_0(X0,X1,X2) != k3_xtuple_0(sK0,sK1,sK2) | sK5 = X2 <- (4) [superposition 146,56] {a:3,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
150. sK2 = sK5 <- (4) [equality resolution 148] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
151. 3 | ~4 [avatar split clause 150,77,72]
154. k3_xtuple_0(sK0,sK1,sK2) = k3_xtuple_0(sK3,sK4,sK2) <- (3, 4) [backward demodulation 79,73] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
156. 10 <=> k3_xtuple_0(sK0,sK1,sK2) = k3_xtuple_0(sK3,sK4,sK2) [avatar definition]
158. k3_xtuple_0(sK0,sK1,sK2) = k3_xtuple_0(sK3,sK4,sK2) <- (10) [avatar component clause 156] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
159. 10 | ~3 | ~4 [avatar split clause 154,77,72,156]
179. k4_tarski(X0,X1) != k3_xtuple_0(sK0,sK1,sK2) | k4_tarski(sK3,sK4) = X0 <- (10) [superposition 121,158] {a:2,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
183. k3_xtuple_0(X0,X1,X2) != k3_xtuple_0(sK0,sK1,sK2) | k4_tarski(X0,X1) = k4_tarski(sK3,sK4) <- (10) [superposition 179,56] {a:3,w:16,wCS:1600,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
207. k4_tarski(sK3,sK4) = k4_tarski(sK0,sK1) <- (10) [equality resolution 183] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
210. 11 <=> k4_tarski(sK3,sK4) = k4_tarski(sK0,sK1) [avatar definition]
212. k4_tarski(sK3,sK4) = k4_tarski(sK0,sK1) <- (11) [avatar component clause 210] {a:4,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
213. 11 | ~10 [avatar split clause 207,156,210]
217. $false <- (~1, 11) [unit resulting resolution 66,212,57] {a:5,w:0,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
227. 1 | ~11 [avatar contradiction clause 217]
235. k4_tarski(sK0,sK1) = k4_tarski(sK0,sK4) <- (1, 11) [backward demodulation 212,65] {a:4,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
238. 13 <=> k4_tarski(sK0,sK1) = k4_tarski(sK0,sK4) [avatar definition]
240. k4_tarski(sK0,sK1) = k4_tarski(sK0,sK4) <- (13) [avatar component clause 238] {a:4,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
241. 13 | ~1 | ~11 [avatar split clause 235,210,64,238]
254. $false <- (~2, 13) [unit resulting resolution 70,240,58] {a:5,w:0,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
264. 2 | ~13 [avatar contradiction clause 254]
267. $false [avatar sat refutation 75,80,151,159,213,227,241,264] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:93,thDist:-93}
% SZS output end Proof for xtuple_0__t3_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 3
% 
% >>> Saturation
% Initial clauses: 16
% Generated clauses: 166
% Active clauses: 38
% Passive clauses: 88
% Final active clauses: 28
% Final passive clauses: 41
% Main loop iterations started: 47
% 
% >>> Simplifying Inferences
% Fw demodulations: 15
% Bw demodulations: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Equational tautologies: 3
% Forward subsumptions: 38
% 
% >>> Generating Inferences
% Unit resulting resolution: 4
% Forward superposition: 40
% Backward superposition: 47
% Self superposition: 7
% Equality resolution: 5
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 3
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 19
% SAT solver unit clauses: 6
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% superposition: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

