i: [1,49,0,2,1,1,0,0] t28_yellow_5
a: [3,76,0,2,1,49]
s: 76
i: [1,48,0,2,1,1,0,0] t28_yellow_5
a: [3,81,0,2,1,48]
s: 81
i: [1,47,0,2,1,1,0,0] t28_yellow_5
a: [3,86,0,2,1,47]
s: 86
i: [1,46,0,2,1,1,0,0] t28_yellow_5
a: [3,91,0,2,1,46]
s: 91
i: [1,50,0,4,1,1,0,0] t28_yellow_5
a: [3,71,0,4,1,50]
s: 71
i: [1,51,0,5,1,1,0,0] t28_yellow_5
a: [3,66,0,5,1,51]
s: 66
i: [1,62,0,2,1,0,0,1] existence_l1_orders_2
a: [3,96,0,2,1,62]
s: 96
i: [1,57,0,4,1,0,0,1] existence_m1_subset_1
s: 57
i: [1,53,0,6,3,0,0,2] cc2_lattice3
s: 53
i: [1,52,0,7,2,0,0,1] dt_k3_yellow_0
s: 52
i: [1,61,0,17,4,0,0,2] dt_k11_lattice3
s: 61
i: [1,56,0,20,6,0,0,1] t25_yellow_5
s: 56
i: [1,58,0,21,6,0,0,2] dt_k12_lattice3
s: 58
i: [1,54,0,23,6,0,0,1] d3_yellow_5
s: 54
i: [1,55,0,23,6,0,0,1] d3_yellow_5
s: 55
i: [1,59,0,23,6,0,0,2] redefinition_k12_lattice3
s: 59
i: [1,60,0,23,6,0,0,2] commutativity_k12_lattice3
s: 60
d: [2,115,1,35,11,0,44,60,56]
d: [2,126,1,25,7,0,18,115]
d: [2,127,1,20,6,0,21,126,52]
s: 127
d: [2,116,1,35,11,0,44,60,56]
d: [2,125,1,25,7,0,18,116]
d: [2,128,1,20,6,0,21,125,52]
s: 128
d: [2,117,1,35,11,0,44,56,60]
d: [2,124,1,25,7,0,18,117]
d: [2,129,1,20,6,0,21,124,52]
s: 129
d: [2,119,1,35,11,0,44,56,60]
d: [2,122,1,25,7,0,18,119]
d: [2,130,1,20,6,0,21,122,52]
s: 130
d: [2,104,1,37,11,0,44,59,55]
d: [2,111,1,27,8,0,18,104]
d: [2,112,1,25,7,0,21,111,53]
s: 112
d: [2,105,1,37,11,0,44,55,59]
d: [2,110,1,27,8,0,18,105]
d: [2,113,1,25,7,0,21,110,53]
s: 113
d: [2,106,1,37,11,0,44,54,59]
d: [2,109,1,27,8,0,18,106]
d: [2,114,1,25,7,0,21,109,53]
s: 114
d: [2,151,2,37,12,0,44,114,56]
d: [2,161,2,27,8,0,18,151]
d: [2,162,2,22,7,0,19,161]
d: [2,163,2,17,6,0,21,162,52]
s: 163
d: [2,165,3,0,0,6,52,81,91,86,76,71,66,163]
e: 165
d: [2,166,3,12,5,1,40,163,66]
d: [2,168,3,8,4,2,21,166,71]
d: [2,169,3,6,3,3,21,168,76]
d: [2,170,3,4,2,4,21,169,86]
d: [2,171,3,2,1,5,21,170,91]
d: [2,172,3,0,0,6,21,171,81]
e: 172
f: 172,46,47,48,49,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for yellow_5__t28_yellow_5
% SZS output start Proof for yellow_5__t28_yellow_5
1. ! [X0] : ((l1_orders_2(X0) & v1_yellow_0(X0) & v2_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ~r1_yellow_5(X0,X1,k3_yellow_0(X0)))) [input t28_yellow_5]
2. ~! [X0] : ((l1_orders_2(X0) & v1_yellow_0(X0) & v2_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ~r1_yellow_5(X0,X1,k3_yellow_0(X0)))) [negated conjecture 1]
3. ! [X0] : (l1_orders_2(X0) => (v2_lattice3(X0) => ~v2_struct_0(X0))) [input cc2_lattice3]
5. ! [X0] : ((l1_orders_2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r1_yellow_5(X0,X1,X2) <=> k3_yellow_0(X0) != k11_lattice3(X0,X1,X2))))) [input d3_yellow_5]
8. ! [X0] : (l1_orders_2(X0) => m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0))) [input dt_k3_yellow_0]
16. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_orders_2(X0) & v2_lattice3(X0) & v5_orders_2(X0)) => k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2)) [input redefinition_k12_lattice3]
17. ! [X0] : ((l1_orders_2(X0) & v1_yellow_0(X0) & v2_lattice3(X0) & v5_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k3_yellow_0(X0) = k12_lattice3(X0,X1,k3_yellow_0(X0)))) [input t25_yellow_5]
20. ~! [X0] : ((l1_orders_2(X0) & v1_yellow_0(X0) & v2_lattice3(X0) & v5_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ~r1_yellow_5(X0,X1,k3_yellow_0(X0)))) [pure predicate removal 2]
21. ? [X0] : (? [X1] : (r1_yellow_5(X0,X1,k3_yellow_0(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_orders_2(X0) & v1_yellow_0(X0) & v2_lattice3(X0) & v5_orders_2(X0))) [ennf transformation 20]
22. ? [X0] : (? [X1] : (r1_yellow_5(X0,X1,k3_yellow_0(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v1_yellow_0(X0) & v2_lattice3(X0) & v5_orders_2(X0)) [flattening 21]
23. ! [X0] : (m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0)) [ennf transformation 8]
24. ! [X0] : ((~v2_struct_0(X0) | ~v2_lattice3(X0)) | ~l1_orders_2(X0)) [ennf transformation 3]
25. ! [X0] : (~v2_struct_0(X0) | ~v2_lattice3(X0) | ~l1_orders_2(X0)) [flattening 24]
26. ! [X0] : (! [X1] : (! [X2] : ((r1_yellow_5(X0,X1,X2) <=> k3_yellow_0(X0) != k11_lattice3(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 5]
27. ! [X0] : (! [X1] : (! [X2] : ((r1_yellow_5(X0,X1,X2) <=> k3_yellow_0(X0) != k11_lattice3(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [flattening 26]
28. ! [X0] : (! [X1] : (k3_yellow_0(X0) = k12_lattice3(X0,X1,k3_yellow_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0))) [ennf transformation 17]
29. ! [X0] : (! [X1] : (k3_yellow_0(X0) = k12_lattice3(X0,X1,k3_yellow_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0)) [flattening 28]
32. ! [X0,X1,X2] : (k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0))) [ennf transformation 16]
33. ! [X0,X1,X2] : (k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0)) [flattening 32]
38. ? [X0] : (? [X1] : (r1_yellow_5(X0,X1,k3_yellow_0(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v1_yellow_0(X0) & v2_lattice3(X0) & v5_orders_2(X0)) => (? [X1] : (r1_yellow_5(sK0,X1,k3_yellow_0(sK0)) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v1_yellow_0(sK0) & v2_lattice3(sK0) & v5_orders_2(sK0)) [choice axiom]
39. ? [X1] : (r1_yellow_5(sK0,X1,k3_yellow_0(sK0)) & m1_subset_1(X1,u1_struct_0(sK0))) => (r1_yellow_5(sK0,sK1,k3_yellow_0(sK0)) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
40. (r1_yellow_5(sK0,sK1,k3_yellow_0(sK0)) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v1_yellow_0(sK0) & v2_lattice3(sK0) & v5_orders_2(sK0) [skolemisation 22,39,38]
41. ! [X0] : (! [X1] : (! [X2] : (((r1_yellow_5(X0,X1,X2) | k3_yellow_0(X0) = k11_lattice3(X0,X1,X2)) & (k3_yellow_0(X0) != k11_lattice3(X0,X1,X2) | ~r1_yellow_5(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [nnf transformation 27]
46. v5_orders_2(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
47. v2_lattice3(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. v1_yellow_0(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. l1_orders_2(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 40] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. r1_yellow_5(sK0,sK1,k3_yellow_0(sK0)) [cnf transformation 40] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 23] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. ~v2_struct_0(X0) | ~v2_lattice3(X0) | ~l1_orders_2(X0) [cnf transformation 25] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
54. k3_yellow_0(X0) != k11_lattice3(X0,X1,X2) | ~r1_yellow_5(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0) [cnf transformation 41] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. k3_yellow_0(X0) = k12_lattice3(X0,X1,k3_yellow_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 29] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 33] {a:0,w:23,wCS:9200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
64. 1 <=> r1_yellow_5(sK0,sK1,k3_yellow_0(sK0)) [avatar definition]
66. r1_yellow_5(sK0,sK1,k3_yellow_0(sK0)) <- (1) [avatar component clause 64] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
67. 1 [avatar split clause 51,64]
69. 2 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
71. m1_subset_1(sK1,u1_struct_0(sK0)) <- (2) [avatar component clause 69] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. 2 [avatar split clause 50,69]
74. 3 <=> l1_orders_2(sK0) [avatar definition]
76. l1_orders_2(sK0) <- (3) [avatar component clause 74] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
77. 3 [avatar split clause 49,74]
79. 4 <=> v1_yellow_0(sK0) [avatar definition]
81. v1_yellow_0(sK0) <- (4) [avatar component clause 79] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. 4 [avatar split clause 48,79]
84. 5 <=> v2_lattice3(sK0) [avatar definition]
86. v2_lattice3(sK0) <- (5) [avatar component clause 84] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
87. 5 [avatar split clause 47,84]
89. 6 <=> v5_orders_2(sK0) [avatar definition]
91. v5_orders_2(sK0) <- (6) [avatar component clause 89] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. 6 [avatar split clause 46,89]
106. k3_yellow_0(X3) != k12_lattice3(X3,X4,X5) | ~r1_yellow_5(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | v2_struct_0(X3) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) [superposition 54,59] {a:1,w:37,wCS:14800,sine:1,thAx:0,allAx:2,thDist:-2}
109. k3_yellow_0(X3) != k12_lattice3(X3,X4,X5) | ~r1_yellow_5(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | v2_struct_0(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) [duplicate literal removal 106] {a:1,w:27,wCS:10800,sine:1,thAx:0,allAx:2,thDist:-2}
114. k3_yellow_0(X3) != k12_lattice3(X3,X4,X5) | ~r1_yellow_5(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) [subsumption resolution 109,53] {a:1,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
151. k3_yellow_0(X0) != k3_yellow_0(X0) | ~r1_yellow_5(X0,X1,k3_yellow_0(X0)) | ~m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) [superposition 114,56] {a:2,w:37,wCS:14800,sine:1,thAx:0,allAx:3,thDist:-3}
161. k3_yellow_0(X0) != k3_yellow_0(X0) | ~r1_yellow_5(X0,X1,k3_yellow_0(X0)) | ~m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v1_yellow_0(X0) [duplicate literal removal 151] {a:2,w:27,wCS:10800,sine:1,thAx:0,allAx:3,thDist:-3}
162. ~r1_yellow_5(X0,X1,k3_yellow_0(X0)) | ~m1_subset_1(k3_yellow_0(X0),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v1_yellow_0(X0) [trivial inequality removal 161] {a:2,w:22,wCS:8800,sine:1,thAx:0,allAx:3,thDist:-3}
163. ~r1_yellow_5(X0,X1,k3_yellow_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) | ~v1_yellow_0(X0) [subsumption resolution 162,52] {a:2,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
166. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) | ~v1_yellow_0(sK0) <- (1) [resolution 163,66] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
168. ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) | ~v1_yellow_0(sK0) <- (1, 2) [subsumption resolution 166,71] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
169. ~v2_lattice3(sK0) | ~v5_orders_2(sK0) | ~v1_yellow_0(sK0) <- (1, 2, 3) [subsumption resolution 168,76] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
170. ~v5_orders_2(sK0) | ~v1_yellow_0(sK0) <- (1, 2, 3, 5) [subsumption resolution 169,86] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
171. ~v1_yellow_0(sK0) <- (1, 2, 3, 5, 6) [subsumption resolution 170,91] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
172. $false <- (1, 2, 3, 4, 5, 6) [subsumption resolution 171,81] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
173. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar contradiction clause 172]
174. $false [avatar sat refutation 67,72,77,82,87,92,173] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:54,thDist:-54}
% SZS output end Proof for yellow_5__t28_yellow_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 98
% Active clauses: 21
% Passive clauses: 32
% Final active clauses: 21
% Final passive clauses: 7
% Main loop iterations started: 25
% 
% >>> Simplifying Inferences
% Duplicate literals: 117
% Trivial inequalities: 2
% Fw subsumption resolutions: 14
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 1
% Unit resulting resolution: 1
% Forward superposition: 13
% Backward superposition: 15
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

