i: [1,62,0,7,2,1,0,0] t3_subset
a: [3,98,0,3,1,62]
s: 98
a: [3,95,0,4,1,62]
s: 95
i: [1,63,0,3,1,0,0,1] fc1_subset_1
s: 63
i: [1,85,0,3,1,0,0,1] rc2_subset_1
s: 85
i: [1,88,0,3,1,0,0,1] reflexivity_r1_tarski
s: 88
i: [1,83,0,4,1,0,0,1] existence_m1_subset_1
s: 83
i: [1,87,0,4,1,0,0,255] rc3_subset_1
s: 87
i: [1,67,0,5,2,0,0,1] rc1_subset_1
s: 67
i: [1,70,0,5,2,0,0,255] t6_boole
s: 70
i: [1,82,0,5,2,0,0,2] t7_boole
s: 82
i: [1,84,0,5,1,0,0,1] rc2_subset_1
s: 84
i: [1,86,0,5,1,0,0,255] rc3_subset_1
s: 86
i: [1,69,0,6,2,0,0,255] rc4_subset_1
s: 69
i: [1,89,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 89
i: [1,66,0,7,2,0,0,1] rc1_subset_1
s: 66
i: [1,68,0,7,2,0,0,255] rc4_subset_1
s: 68
i: [1,75,0,7,3,0,0,2] d1_subset_1
s: 75
i: [1,76,0,7,3,0,0,2] d1_subset_1
s: 76
d: [2,105,1,5,2,0,40,75,83]
s: 105
i: [1,81,0,7,3,0,0,2] t8_boole
s: 81
i: [1,90,0,7,2,0,0,1] d1_zfmisc_1
s: 90
i: [1,91,0,7,2,0,0,1] d1_zfmisc_1
s: 91
i: [1,71,0,8,3,0,0,1] cc1_subset_1
s: 71
d: [2,119,1,5,2,0,40,71,86]
s: 119
i: [1,73,0,8,3,0,0,2] d1_subset_1
s: 73
i: [1,74,0,8,3,0,0,2] d1_subset_1
s: 74
d: [2,124,1,8,2,0,40,73,84]
d: [2,128,1,5,1,0,21,124,63]
s: 128
d: [2,130,0,6,2,0,21,74,82]
s: 130
d: [2,131,2,4,1,0,40,128,91]
s: 131
i: [1,72,0,9,3,0,0,255] cc4_subset_1
s: 72
i: [1,64,0,11,4,0,0,255] cc3_subset_1
s: 64
d: [2,125,1,8,2,0,40,73,86]
d: [2,129,1,5,1,0,21,125,63]
s: 129
i: [1,65,0,11,4,0,0,255] cc2_subset_1
s: 65
d: [2,153,2,4,1,0,40,129,91]
s: 153
i: [1,79,0,14,3,0,0,1] d1_zfmisc_1
s: 79
i: [1,80,0,14,3,0,0,1] d1_zfmisc_1
s: 80
d: [2,147,1,9,3,0,40,64,86]
d: [2,152,1,5,2,0,21,147,87]
s: 152
d: [2,117,1,6,2,0,40,71,83]
s: 117
d: [2,132,2,5,1,0,40,128,89]
s: 132
d: [2,123,1,6,2,0,40,73,83]
s: 123
d: [2,140,1,6,2,0,40,72,84]
s: 140
d: [2,155,2,5,1,0,40,129,89]
s: 155
d: [2,146,1,9,3,0,40,64,84]
d: [2,151,1,6,2,0,21,146,85]
s: 151
d: [2,168,2,8,2,0,40,123,91]
d: [2,169,2,5,1,0,21,168,63]
s: 169
d: [2,102,1,7,3,0,44,70,70]
s: 102
d: [2,111,1,7,2,0,40,90,89]
s: 111
d: [2,166,2,6,2,0,40,123,89]
s: 166
d: [2,121,1,10,3,0,40,73,66]
d: [2,126,1,7,2,0,21,121,63]
s: 126
d: [2,172,2,6,2,0,40,126,91]
s: 172
d: [2,122,1,10,3,0,40,73,68]
d: [2,127,1,7,2,0,21,122,63]
s: 127
d: [2,134,1,7,2,0,40,130,90]
s: 134
d: [2,180,2,0,0,2,52,98,95,134]
e: 180
d: [2,181,2,3,1,1,40,134,95]
d: [2,188,2,0,0,2,21,181,98]
e: 188
r: 95
r: 98
a: [3,99,0,3,1,62]
s: 99
a: [3,94,0,4,1,62]
s: 94
d: [2,190,1,7,3,1,40,94,64]
a: [3,197,1,2,1,190]
s: 197
a: [3,206,1,2,1,190]
s: 206
d: [2,193,1,7,2,1,40,94,73]
d: [2,210,1,4,1,1,21,193,63]
a: [3,214,1,4,1,210]
s: 214
d: [2,216,2,0,0,2,52,99,214,91]
e: 216
d: [2,217,2,3,1,1,40,214,91]
d: [2,222,2,0,0,2,21,217,99]
e: 222
i: [1,61,0,7,2,1,0,0] t3_subset
f: 222,210,188,61,62
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset__t3_subset
% SZS output start Proof for subset__t3_subset
1. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
2. ~! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [negated conjecture 1]
8. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
9. ! [X0] : ! [X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [input d1_zfmisc_1]
14. ! [X0] : ~v1_xboole_0(k1_zfmisc_1(X0)) [input fc1_subset_1]
21. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
23. ~! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 2]
24. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 8]
25. ! [X0,X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [flattening 9]
27. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 21]
29. ? [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <~> r1_tarski(X0,X1)) [ennf transformation 23]
39. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 24]
41. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 27]
43. ? [X0,X1] : ((~r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1))) & (r1_tarski(X0,X1) | m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 29]
44. ? [X0,X1] : ((~r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1))) & (r1_tarski(X0,X1) | m1_subset_1(X0,k1_zfmisc_1(X1)))) => ((~r1_tarski(sK0,sK1) | ~m1_subset_1(sK0,k1_zfmisc_1(sK1))) & (r1_tarski(sK0,sK1) | m1_subset_1(sK0,k1_zfmisc_1(sK1)))) [choice axiom]
45. (~r1_tarski(sK0,sK1) | ~m1_subset_1(sK0,k1_zfmisc_1(sK1))) & (r1_tarski(sK0,sK1) | m1_subset_1(sK0,k1_zfmisc_1(sK1))) [skolemisation 43,44]
50. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 39]
51. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ~r1_tarski(X2,X0)) & (r1_tarski(X2,X0) | ~r2_hidden(X2,X1))) | k1_zfmisc_1(X0) != X1)) [nnf transformation 25]
52. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [rectify 51]
53. ! [X1,X0] : (? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1))) => ((~r1_tarski(sK4(X0,X1),X0) | ~r2_hidden(sK4(X0,X1),X1)) & (r1_tarski(sK4(X0,X1),X0) | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
54. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ((~r1_tarski(sK4(X0,X1),X0) | ~r2_hidden(sK4(X0,X1),X1)) & (r1_tarski(sK4(X0,X1),X0) | r2_hidden(sK4(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [skolemisation 52,53]
61. r1_tarski(sK0,sK1) | m1_subset_1(sK0,k1_zfmisc_1(sK1)) [cnf transformation 45] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
62. ~r1_tarski(sK0,sK1) | ~m1_subset_1(sK0,k1_zfmisc_1(sK1)) [cnf transformation 45] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
63. ~v1_xboole_0(k1_zfmisc_1(X0)) [cnf transformation 14] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. ~m1_subset_1(X1,X0) | r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 50] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
74. m1_subset_1(X1,X0) | ~r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 50] {a:0,w:8,wCS:3200,sine:2,thAx:0,allAx:1,thDist:-1}
77. r1_tarski(X3,X0) | ~r2_hidden(X3,X1) | k1_zfmisc_1(X0) != X1 [cnf transformation 54] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
78. r2_hidden(X3,X1) | ~r1_tarski(X3,X0) | k1_zfmisc_1(X0) != X1 [cnf transformation 54] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
82. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 41] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
90. r2_hidden(X3,k1_zfmisc_1(X0)) | ~r1_tarski(X3,X0) [equality resolution 78] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
91. ~r2_hidden(X3,k1_zfmisc_1(X0)) | r1_tarski(X3,X0) [equality resolution 77] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
93. 1 <=> m1_subset_1(sK0,k1_zfmisc_1(sK1)) [avatar definition]
94. m1_subset_1(sK0,k1_zfmisc_1(sK1)) <- (1) [avatar component clause 93] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
95. ~m1_subset_1(sK0,k1_zfmisc_1(sK1)) <- (~1) [avatar component clause 93] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
97. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
98. r1_tarski(sK0,sK1) <- (2) [avatar component clause 97] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
99. ~r1_tarski(sK0,sK1) <- (~2) [avatar component clause 97] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
100. ~1 | ~2 [avatar split clause 62,97,93]
101. 1 | 2 [avatar split clause 61,97,93]
130. ~r2_hidden(X1,X0) | m1_subset_1(X1,X0) [subsumption resolution 74,82] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
134. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [resolution 130,90] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
181. ~r1_tarski(sK0,sK1) <- (~1) [resolution 134,95] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
188. $false <- (~1, 2) [subsumption resolution 181,98] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
189. 1 | ~2 [avatar contradiction clause 188]
193. r2_hidden(sK0,k1_zfmisc_1(sK1)) | v1_xboole_0(k1_zfmisc_1(sK1)) <- (1) [resolution 94,73] {a:1,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
210. r2_hidden(sK0,k1_zfmisc_1(sK1)) <- (1) [subsumption resolution 193,63] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
212. 6 <=> r2_hidden(sK0,k1_zfmisc_1(sK1)) [avatar definition]
214. r2_hidden(sK0,k1_zfmisc_1(sK1)) <- (6) [avatar component clause 212] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
215. 6 | ~1 [avatar split clause 210,93,212]
217. r1_tarski(sK0,sK1) <- (6) [resolution 214,91] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
222. $false <- (~2, 6) [subsumption resolution 217,99] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
223. 2 | ~6 [avatar contradiction clause 222]
224. $false [avatar sat refutation 100,101,189,215,223] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:56,thDist:-56}
% SZS output end Proof for subset__t3_subset
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 137
% Active clauses: 53
% Passive clauses: 70
% Final active clauses: 51
% Final passive clauses: 14
% Main loop iterations started: 56
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Fw subsumption resolutions: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 12
% Forward subsumptions: 28
% 
% >>> Generating Inferences
% Binary resolution: 79
% Unit resulting resolution: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 11
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

