i: [1,99,0,2,1,1,0,0] t19_robbins2
a: [3,154,0,2,1,99]
s: 154
i: [1,98,0,2,1,1,0,0] t19_robbins2
a: [3,159,0,2,1,98]
s: 159
i: [1,97,0,2,1,1,0,0] t19_robbins2
a: [3,164,0,2,1,97]
s: 164
i: [1,101,0,4,1,1,0,0] t19_robbins2
a: [3,144,0,4,1,101]
s: 144
i: [1,100,0,4,1,1,0,0] t19_robbins2
a: [3,149,0,4,1,100]
s: 149
i: [1,132,0,2,1,0,0,2] existence_l2_lattices
a: [3,169,0,2,1,132]
s: 169
i: [1,133,0,2,1,0,0,2] existence_l1_robbins1
a: [3,174,0,2,1,133]
s: 174
i: [1,134,0,2,1,0,0,1] existence_l2_robbins1
a: [3,179,0,2,1,134]
s: 179
i: [1,135,0,2,1,0,0,3] existence_l1_struct_0
a: [3,184,0,2,1,135]
s: 184
i: [1,102,0,15,1,1,0,0] t19_robbins2
a: [3,139,0,15,1,102]
s: 139
i: [1,103,0,4,2,0,0,2] dt_l2_lattices
s: 103
d: [2,186,1,2,1,1,40,103,169]
a: [3,190,1,2,1,186]
s: 190
i: [1,104,0,4,2,0,0,2] dt_l1_robbins1
s: 104
d: [2,192,1,2,1,1,40,104,174]
a: [3,196,1,2,1,192]
s: 196
i: [1,105,0,4,2,0,0,2] dt_l2_robbins1
s: 105
i: [1,106,0,4,2,0,0,2] dt_l2_robbins1
s: 106
d: [2,198,1,2,1,1,40,105,154]
a: [3,203,1,2,1,198]
s: 203
i: [1,129,0,4,1,0,0,1] existence_m1_subset_1
s: 129
d: [2,210,1,2,1,1,40,106,154]
a: [3,215,1,2,1,210]
s: 215
i: [1,107,0,6,3,0,0,255] cc7_struct_0
s: 107
i: [1,108,0,6,3,0,0,255] cc5_struct_0
s: 108
d: [2,222,2,2,1,1,40,203,103]
a: [3,226,2,2,1,222]
s: 226
i: [1,110,0,6,3,0,0,255] cc17_struct_0
s: 110
d: [2,199,1,2,1,1,40,105,179]
a: [3,208,1,2,1,199]
s: 208
i: [1,111,0,6,3,0,0,255] cc2_struct_0
s: 111
i: [1,112,0,6,3,0,0,255] cc6_struct_0
s: 112
d: [2,211,1,2,1,1,40,106,179]
a: [3,220,1,2,1,211]
s: 220
i: [1,113,0,6,3,0,0,255] cc1_struct_0
s: 113
d: [2,230,1,4,2,1,40,108,184]
a: [3,236,1,2,1,230]
s: 236
i: [1,116,0,6,3,0,0,255] cc4_struct_0
s: 116
i: [1,114,0,7,3,0,0,255] cc8_struct_0
s: 114
d: [2,231,2,4,2,1,40,108,190]
a: [3,245,2,2,1,231]
s: 245
i: [1,117,0,7,3,0,0,255] cc9_struct_0
s: 117
d: [2,232,2,4,2,1,40,108,196]
a: [3,254,2,2,1,232]
s: 254
i: [1,118,0,7,3,0,0,255] cc11_struct_0
s: 118
i: [1,119,0,7,3,0,0,255] cc11_struct_0
s: 119
d: [2,280,2,2,1,1,40,208,103]
a: [3,284,2,2,1,280]
s: 284
i: [1,121,0,7,3,0,0,255] fc9_struct_0
s: 121
d: [2,290,3,4,2,1,40,284,110]
a: [3,295,3,2,1,290]
s: 295
i: [1,122,0,7,3,0,0,255] fc6_struct_0
s: 122
i: [1,125,0,7,3,0,0,255] fc2_struct_0
s: 125
i: [1,126,0,7,3,0,0,255] fc8_struct_0
s: 126
i: [1,127,0,7,3,0,0,255] fc7_struct_0
s: 127
i: [1,128,0,7,3,0,0,255] fc1_struct_0
s: 128
i: [1,120,0,9,4,0,0,255] cc10_struct_0
s: 120
i: [1,130,0,14,4,0,0,1] dt_k3_robbins1
s: 130
i: [1,131,0,19,5,0,0,1] dt_k1_lattices
s: 131
i: [1,123,0,27,6,0,0,1] t14_robbins2
s: 123
d: [2,316,1,36,6,1,44,139,123]
d: [2,322,1,34,5,2,21,316,154]
d: [2,323,1,32,4,3,21,322,159]
d: [2,324,1,30,3,4,21,323,164]
a: [3,328,1,9,1,324]
s: 328
i: [1,124,0,29,6,0,0,1] t18_robbins2
s: 124
d: [2,356,1,43,11,0,44,124,123]
d: [2,361,1,29,6,0,18,356]
s: 361
d: [2,344,2,11,3,1,40,328,130]
d: [2,345,2,9,2,2,21,344,215]
d: [2,346,2,7,1,3,21,345,164]
a: [3,350,2,7,1,346]
s: 350
d: [2,388,3,0,0,5,52,164,203,149,144,350,131]
e: 388
d: [2,389,3,12,4,1,40,350,131]
d: [2,391,3,8,3,2,21,389,144]
d: [2,392,3,4,2,3,21,391,149]
d: [2,393,3,2,1,4,21,392,203]
d: [2,394,3,0,0,5,21,393,164]
e: 394
r: 328
r: 350
a: [3,332,1,6,1,324]
s: 332
d: [2,396,2,0,0,4,52,164,215,144,332,130]
e: 396
d: [2,397,2,8,3,1,40,332,130]
d: [2,399,2,4,2,2,21,397,144]
d: [2,400,2,2,1,3,21,399,215]
d: [2,401,2,0,0,4,21,400,164]
e: 401
r: 332
a: [3,331,1,6,1,324]
s: 331
a: [3,327,1,9,1,324]
s: 327
a: [3,336,1,15,1,324]
s: 336
d: [2,403,2,0,0,6,52,164,159,154,144,149,336,124]
e: 403
d: [2,404,2,17,6,1,44,336,124]
d: [2,405,2,14,5,1,19,404]
d: [2,407,2,10,4,2,21,405,149]
d: [2,408,2,6,3,3,21,407,144]
d: [2,409,2,4,2,4,21,408,154]
d: [2,410,2,2,1,5,21,409,159]
d: [2,411,2,0,0,6,21,410,164]
e: 411
d: [2,318,1,36,6,1,44,139,123]
d: [2,339,1,34,5,2,21,318,154]
d: [2,340,1,32,4,3,21,339,159]
d: [2,341,1,30,3,4,21,340,164]
f: 411,401,394,346,341,210,198,97,98,99,100,101,102
% Refutation found. Thanks to Tanya!
% SZS status Theorem for robbins2__t19_robbins2
% SZS output start Proof for robbins2__t19_robbins2
1. ! [X0] : ((l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,k1_lattices(X0,X1,X2)),k3_robbins1(X0,X2))) = X2))) [input t19_robbins2]
2. ~! [X0] : ((l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,k1_lattices(X0,X1,X2)),k3_robbins1(X0,X2))) = X2))) [negated conjecture 1]
14. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l2_lattices(X0) & ~v2_struct_0(X0)) => m1_subset_1(k1_lattices(X0,X1,X2),u1_struct_0(X0))) [input dt_k1_lattices]
16. ! [X0,X1] : ((m1_subset_1(X1,u1_struct_0(X0)) & l1_robbins1(X0) & ~v2_struct_0(X0)) => m1_subset_1(k3_robbins1(X0,X1),u1_struct_0(X0))) [input dt_k3_robbins1]
20. ! [X0] : (l2_robbins1(X0) => (l1_robbins1(X0) & l2_lattices(X0))) [input dt_l2_robbins1]
34. ! [X0] : ((l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => k3_robbins1(X0,k1_lattices(X0,X1,X2)) = k3_robbins1(X0,k1_lattices(X0,X2,X1))))) [input t14_robbins2]
35. ! [X0] : ((l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,X1),k3_robbins1(X0,k1_lattices(X0,X2,X1)))) = X1))) [input t18_robbins2]
36. ? [X0] : (? [X1] : (? [X2] : (k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,k1_lattices(X0,X1,X2)),k3_robbins1(X0,X2))) != X2 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
37. ? [X0] : (? [X1] : (? [X2] : (k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,k1_lattices(X0,X1,X2)),k3_robbins1(X0,X2))) != X2 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) [flattening 36]
40. ! [X0] : ((l1_robbins1(X0) & l2_lattices(X0)) | ~l2_robbins1(X0)) [ennf transformation 20]
67. ! [X0] : (! [X1] : (! [X2] : (k3_robbins1(X0,k1_lattices(X0,X1,X2)) = k3_robbins1(X0,k1_lattices(X0,X2,X1)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0))) [ennf transformation 34]
68. ! [X0] : (! [X1] : (! [X2] : (k3_robbins1(X0,k1_lattices(X0,X1,X2)) = k3_robbins1(X0,k1_lattices(X0,X2,X1)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0)) [flattening 67]
69. ! [X0] : (! [X1] : (! [X2] : (k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,X1),k3_robbins1(X0,k1_lattices(X0,X2,X1)))) = X1 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0))) [ennf transformation 35]
70. ! [X0] : (! [X1] : (! [X2] : (k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,X1),k3_robbins1(X0,k1_lattices(X0,X2,X1)))) = X1 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0)) [flattening 69]
79. ! [X0,X1] : (m1_subset_1(k3_robbins1(X0,X1),u1_struct_0(X0)) | (~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_robbins1(X0) | v2_struct_0(X0))) [ennf transformation 16]
80. ! [X0,X1] : (m1_subset_1(k3_robbins1(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_robbins1(X0) | v2_struct_0(X0)) [flattening 79]
81. ! [X0,X1,X2] : (m1_subset_1(k1_lattices(X0,X1,X2),u1_struct_0(X0)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | v2_struct_0(X0))) [ennf transformation 14]
82. ! [X0,X1,X2] : (m1_subset_1(k1_lattices(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | v2_struct_0(X0)) [flattening 81]
83. ? [X0] : (? [X1] : (? [X2] : (k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,k1_lattices(X0,X1,X2)),k3_robbins1(X0,X2))) != X2 & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l2_robbins1(X0) & v1_robbins2(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,k1_lattices(sK0,X1,X2)),k3_robbins1(sK0,X2))) != X2 & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l2_robbins1(sK0) & v1_robbins2(sK0) & ~v2_struct_0(sK0)) [choice axiom]
84. ? [X1] : (? [X2] : (k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,k1_lattices(sK0,X1,X2)),k3_robbins1(sK0,X2))) != X2 & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,k1_lattices(sK0,sK1,X2)),k3_robbins1(sK0,X2))) != X2 & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
85. ? [X2] : (k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,k1_lattices(sK0,sK1,X2)),k3_robbins1(sK0,X2))) != X2 & m1_subset_1(X2,u1_struct_0(sK0))) => (sK2 != k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),k3_robbins1(sK0,sK2))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
86. ((sK2 != k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),k3_robbins1(sK0,sK2))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l2_robbins1(sK0) & v1_robbins2(sK0) & ~v2_struct_0(sK0) [skolemisation 37,85,84,83]
97. ~v2_struct_0(sK0) [cnf transformation 86] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
98. v1_robbins2(sK0) [cnf transformation 86] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
99. l2_robbins1(sK0) [cnf transformation 86] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
100. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 86] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
101. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 86] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
102. sK2 != k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),k3_robbins1(sK0,sK2))) [cnf transformation 86] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
105. ~l2_robbins1(X0) | l2_lattices(X0) [cnf transformation 40] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
106. ~l2_robbins1(X0) | l1_robbins1(X0) [cnf transformation 40] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
123. k3_robbins1(X0,k1_lattices(X0,X1,X2)) = k3_robbins1(X0,k1_lattices(X0,X2,X1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0) [cnf transformation 68] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
124. k3_robbins1(X0,k1_lattices(X0,k3_robbins1(X0,X1),k3_robbins1(X0,k1_lattices(X0,X2,X1)))) = X1 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_robbins1(X0) | ~v1_robbins2(X0) | v2_struct_0(X0) [cnf transformation 70] {a:0,w:29,wCS:11600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
130. m1_subset_1(k3_robbins1(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_robbins1(X0) | v2_struct_0(X0) [cnf transformation 80] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
131. m1_subset_1(k1_lattices(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | v2_struct_0(X0) [cnf transformation 82] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
137. 1 <=> sK2 = k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),k3_robbins1(sK0,sK2))) [avatar definition]
139. sK2 != k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),k3_robbins1(sK0,sK2))) <- (~1) [avatar component clause 137] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
140. ~1 [avatar split clause 102,137]
142. 2 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
144. m1_subset_1(sK2,u1_struct_0(sK0)) <- (2) [avatar component clause 142] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
145. 2 [avatar split clause 101,142]
147. 3 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
149. m1_subset_1(sK1,u1_struct_0(sK0)) <- (3) [avatar component clause 147] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
150. 3 [avatar split clause 100,147]
152. 4 <=> l2_robbins1(sK0) [avatar definition]
154. l2_robbins1(sK0) <- (4) [avatar component clause 152] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
155. 4 [avatar split clause 99,152]
157. 5 <=> v1_robbins2(sK0) [avatar definition]
159. v1_robbins2(sK0) <- (5) [avatar component clause 157] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
160. 5 [avatar split clause 98,157]
162. 6 <=> v2_struct_0(sK0) [avatar definition]
164. ~v2_struct_0(sK0) <- (~6) [avatar component clause 162] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
165. ~6 [avatar split clause 97,162]
198. l2_lattices(sK0) <- (4) [resolution 105,154] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
201. 13 <=> l2_lattices(sK0) [avatar definition]
203. l2_lattices(sK0) <- (13) [avatar component clause 201] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
204. 13 | ~4 [avatar split clause 198,152,201]
210. l1_robbins1(sK0) <- (4) [resolution 106,154] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
213. 15 <=> l1_robbins1(sK0) [avatar definition]
215. l1_robbins1(sK0) <- (15) [avatar component clause 213] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
216. 15 | ~4 [avatar split clause 210,152,213]
318. sK2 != k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,sK2),k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)))) | ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),u1_struct_0(sK0)) | ~m1_subset_1(k3_robbins1(sK0,sK2),u1_struct_0(sK0)) | ~l2_robbins1(sK0) | ~v1_robbins2(sK0) | v2_struct_0(sK0) <- (~1) [superposition 139,123] {a:1,w:36,wCS:3600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
326. 31 <=> m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),u1_struct_0(sK0)) [avatar definition]
328. ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),u1_struct_0(sK0)) <- (~31) [avatar component clause 326] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
330. 32 <=> m1_subset_1(k3_robbins1(sK0,sK2),u1_struct_0(sK0)) [avatar definition]
332. ~m1_subset_1(k3_robbins1(sK0,sK2),u1_struct_0(sK0)) <- (~32) [avatar component clause 330] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
334. 33 <=> sK2 = k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,sK2),k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)))) [avatar definition]
336. sK2 != k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,sK2),k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)))) <- (~33) [avatar component clause 334] {a:1,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
339. sK2 != k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,sK2),k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)))) | ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),u1_struct_0(sK0)) | ~m1_subset_1(k3_robbins1(sK0,sK2),u1_struct_0(sK0)) | ~v1_robbins2(sK0) | v2_struct_0(sK0) <- (~1, 4) [subsumption resolution 318,154] {a:1,w:34,wCS:3400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
340. sK2 != k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,sK2),k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)))) | ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),u1_struct_0(sK0)) | ~m1_subset_1(k3_robbins1(sK0,sK2),u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~1, 4, 5) [subsumption resolution 339,159] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
341. sK2 != k3_robbins1(sK0,k1_lattices(sK0,k3_robbins1(sK0,sK2),k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)))) | ~m1_subset_1(k3_robbins1(sK0,k1_lattices(sK0,sK1,sK2)),u1_struct_0(sK0)) | ~m1_subset_1(k3_robbins1(sK0,sK2),u1_struct_0(sK0)) <- (~1, 4, 5, ~6) [subsumption resolution 340,164] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
342. ~32 | ~31 | ~33 | 1 | ~4 | ~5 | 6 [avatar split clause 341,162,157,152,137,334,326,330]
344. ~m1_subset_1(k1_lattices(sK0,sK1,sK2),u1_struct_0(sK0)) | ~l1_robbins1(sK0) | v2_struct_0(sK0) <- (~31) [resolution 328,130] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
345. ~m1_subset_1(k1_lattices(sK0,sK1,sK2),u1_struct_0(sK0)) | v2_struct_0(sK0) <- (15, ~31) [subsumption resolution 344,215] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
346. ~m1_subset_1(k1_lattices(sK0,sK1,sK2),u1_struct_0(sK0)) <- (~6, 15, ~31) [subsumption resolution 345,164] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
348. 34 <=> m1_subset_1(k1_lattices(sK0,sK1,sK2),u1_struct_0(sK0)) [avatar definition]
350. ~m1_subset_1(k1_lattices(sK0,sK1,sK2),u1_struct_0(sK0)) <- (~34) [avatar component clause 348] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
351. ~34 | 6 | ~15 | 31 [avatar split clause 346,326,213,162,348]
389. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l2_lattices(sK0) | v2_struct_0(sK0) <- (~34) [resolution 350,131] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
391. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l2_lattices(sK0) | v2_struct_0(sK0) <- (2, ~34) [subsumption resolution 389,144] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
392. ~l2_lattices(sK0) | v2_struct_0(sK0) <- (2, 3, ~34) [subsumption resolution 391,149] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
393. v2_struct_0(sK0) <- (2, 3, 13, ~34) [subsumption resolution 392,203] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
394. $false <- (2, 3, ~6, 13, ~34) [subsumption resolution 393,164] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
395. ~2 | ~3 | 6 | ~13 | 34 [avatar contradiction clause 394]
397. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_robbins1(sK0) | v2_struct_0(sK0) <- (~32) [resolution 332,130] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
399. ~l1_robbins1(sK0) | v2_struct_0(sK0) <- (2, ~32) [subsumption resolution 397,144] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
400. v2_struct_0(sK0) <- (2, 15, ~32) [subsumption resolution 399,215] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
401. $false <- (2, ~6, 15, ~32) [subsumption resolution 400,164] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
402. ~2 | 6 | ~15 | 32 [avatar contradiction clause 401]
404. sK2 != sK2 | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_robbins1(sK0) | ~v1_robbins2(sK0) | v2_struct_0(sK0) <- (~33) [superposition 336,124] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
405. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_robbins1(sK0) | ~v1_robbins2(sK0) | v2_struct_0(sK0) <- (~33) [trivial inequality removal 404] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
407. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_robbins1(sK0) | ~v1_robbins2(sK0) | v2_struct_0(sK0) <- (3, ~33) [subsumption resolution 405,149] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
408. ~l2_robbins1(sK0) | ~v1_robbins2(sK0) | v2_struct_0(sK0) <- (2, 3, ~33) [subsumption resolution 407,144] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
409. ~v1_robbins2(sK0) | v2_struct_0(sK0) <- (2, 3, 4, ~33) [subsumption resolution 408,154] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
410. v2_struct_0(sK0) <- (2, 3, 4, 5, ~33) [subsumption resolution 409,159] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
411. $false <- (2, 3, 4, 5, ~6, ~33) [subsumption resolution 410,164] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
412. ~2 | ~3 | ~4 | ~5 | 6 | 33 [avatar contradiction clause 411]
413. $false [avatar sat refutation 140,145,150,155,160,165,204,216,342,351,395,402,412] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:282,thDist:-282}
% SZS output end Proof for robbins2__t19_robbins2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 34
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 39
% Generated clauses: 172
% Active clauses: 52
% Passive clauses: 72
% Final active clauses: 49
% Final passive clauses: 16
% Main loop iterations started: 56
% 
% >>> Simplifying Inferences
% Duplicate literals: 66
% Trivial inequalities: 1
% Fw subsumption resolutions: 22
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 29
% Unit resulting resolution: 3
% Forward superposition: 11
% Backward superposition: 8
% Self superposition: 4
% 
% >>> AVATAR
% Split clauses: 10
% Split components: 22
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 39
% SAT solver unit clauses: 10
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

