i: [1,64,0,3,1,1,0,0] t4_subset_1
a: [3,96,0,3,1,64]
s: 96
i: [1,63,0,4,1,1,0,0] t4_subset_1
a: [3,101,0,4,1,63]
s: 101
i: [1,65,0,6,2,1,0,0] t4_subset_1
s: 65
i: [1,66,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,106,0,2,1,66]
s: 106
i: [1,67,0,2,1,0,0,1] fc1_xboole_0
a: [3,111,0,2,1,67]
s: 111
i: [1,88,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,121,0,2,1,88]
s: 121
i: [1,89,0,2,1,0,0,2] rc2_xboole_0
a: [3,126,0,2,1,89]
s: 126
i: [1,90,0,2,1,0,0,2] rc1_xboole_0
a: [3,131,0,2,1,90]
s: 131
i: [1,92,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,136,0,2,1,92]
s: 136
i: [1,91,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,141,0,2,1,91]
s: 141
i: [1,69,0,3,1,0,0,1] fc1_subset_1
s: 69
i: [1,86,0,3,1,0,0,1] rc2_subset_1
s: 86
i: [1,68,0,3,1,0,0,255] d2_xboole_0
a: [3,116,0,3,1,68]
s: 116
r: 106
i: [1,70,0,4,2,0,0,255] cc2_zfmisc_1
s: 70
i: [1,73,0,4,2,0,0,255] cc1_zfmisc_1
s: 73
i: [1,84,0,4,1,0,0,1] existence_m1_subset_1
s: 84
d: [2,144,1,2,1,1,40,70,111]
a: [3,150,1,2,1,144]
s: 150
i: [1,72,0,5,2,0,0,1] rc1_subset_1
s: 72
d: [2,146,1,2,1,1,40,70,131]
a: [3,155,1,2,1,146]
s: 155
i: [1,74,0,5,2,0,0,1] t6_boole
s: 74
i: [1,75,0,5,2,0,0,1] d1_xboole_0
s: 75
d: [2,145,1,3,1,0,40,70,86]
s: 145
i: [1,83,0,5,2,0,0,1] t7_boole
s: 83
i: [1,85,0,5,1,0,0,1] rc2_subset_1
s: 85
i: [1,76,0,6,2,0,0,1] d1_xboole_0
s: 76
i: [1,87,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 87
d: [2,160,1,6,2,0,40,76,65]
a: [3,165,1,4,1,160]
s: 165
i: [1,71,0,7,2,0,0,1] rc1_subset_1
s: 71
d: [2,171,1,6,2,0,40,87,76]
s: 171
i: [1,79,0,7,3,0,0,1] d1_subset_1
s: 79
i: [1,80,0,7,3,0,0,1] d1_subset_1
s: 80
d: [2,176,2,5,2,1,40,80,165]
a: [3,182,2,2,1,176]
s: 182
i: [1,82,0,7,3,0,0,1] t8_boole
s: 82
d: [2,174,1,5,2,0,40,79,84]
s: 174
i: [1,77,0,8,3,0,0,1] d1_subset_1
s: 77
i: [1,78,0,8,3,0,0,1] d1_subset_1
s: 78
d: [2,190,1,7,2,1,40,77,101]
d: [2,194,1,4,1,1,21,190,69]
a: [3,198,1,4,1,194]
s: 198
d: [2,202,0,6,2,0,21,78,75]
s: 202
d: [2,203,2,4,1,1,40,198,87]
a: [3,208,2,4,1,203]
s: 208
i: [1,81,0,10,3,0,0,1] l2_subset_1
s: 81
d: [2,213,1,6,2,1,40,81,101]
s: 213
d: [2,217,2,6,2,1,40,213,76]
a: [3,221,2,4,1,217]
s: 221
d: [2,223,3,0,0,2,52,165,221,202]
e: 223
d: [2,224,3,4,1,1,40,221,202]
d: [2,228,3,0,0,2,21,224,165]
e: 228
r: 221
a: [3,169,1,2,1,160]
s: 169
d: [2,230,2,0,0,2,52,96,169,74]
e: 230
d: [2,231,2,0,0,3,52,111,96,169,82]
e: 231
d: [2,232,2,0,0,3,52,111,96,169,82]
e: 232
f: 230,228,217,160,63,64
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t4_subset_1
% SZS output start Proof for subset_1__t4_subset_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ~(! [X2] : (m1_subset_1(X2,X0) => ~r2_hidden(X2,X1)) & k1_xboole_0 != X1)) [input t4_subset_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ~(! [X2] : (m1_subset_1(X2,X0) => ~r2_hidden(X2,X1)) & k1_xboole_0 != X1)) [negated conjecture 1]
6. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
7. ! [X0] : (v1_xboole_0(X0) <=> ! [X1] : ~r2_hidden(X1,X0)) [input d1_xboole_0]
16. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (r2_hidden(X2,X1) => r2_hidden(X2,X0))) [input l2_subset_1]
23. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
26. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ~(! [X2] : (m1_subset_1(X2,X0) => ~r2_hidden(X2,X1)) & k1_xboole_0 != X1)) [flattening 2]
27. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 6]
28. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (r2_hidden(X2,X1) => r2_hidden(X2,X0))) [flattening 16]
31. ? [X0,X1] : ((! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,X0)) & k1_xboole_0 != X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 26]
32. ? [X0,X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,X0)) & k1_xboole_0 != X1 & m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 31]
36. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 23]
37. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 27]
38. ! [X0,X1] : (! [X2] : (r2_hidden(X2,X0) | ~r2_hidden(X2,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 28]
42. ? [X0,X1] : (! [X2] : (~r2_hidden(X2,X1) | ~m1_subset_1(X2,X0)) & k1_xboole_0 != X1 & m1_subset_1(X1,k1_zfmisc_1(X0))) => (! [X2] : (~r2_hidden(X2,sK1) | ~m1_subset_1(X2,sK0)) & k1_xboole_0 != sK1 & m1_subset_1(sK1,k1_zfmisc_1(sK0))) [choice axiom]
43. ! [X2] : (~r2_hidden(X2,sK1) | ~m1_subset_1(X2,sK0)) & k1_xboole_0 != sK1 & m1_subset_1(sK1,k1_zfmisc_1(sK0)) [skolemisation 32,42]
46. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X1] : ~r2_hidden(X1,X0) | ~v1_xboole_0(X0))) [nnf transformation 7]
47. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [rectify 46]
48. ! [X0] : (? [X1] : r2_hidden(X1,X0) => r2_hidden(sK3(X0),X0)) [choice axiom]
49. ! [X0] : ((v1_xboole_0(X0) | r2_hidden(sK3(X0),X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [skolemisation 47,48]
50. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 37]
63. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 43] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
64. k1_xboole_0 != sK1 [cnf transformation 43] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. ~r2_hidden(X2,sK1) | ~m1_subset_1(X2,sK0) [cnf transformation 43] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
74. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 36] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
75. ~r2_hidden(X2,X0) | ~v1_xboole_0(X0) [cnf transformation 49] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
76. r2_hidden(sK3(X0),X0) | v1_xboole_0(X0) [cnf transformation 49] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
78. m1_subset_1(X1,X0) | ~r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 50] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
81. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | ~r2_hidden(X2,X1) | r2_hidden(X2,X0) [cnf transformation 38] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
94. 1 <=> k1_xboole_0 = sK1 [avatar definition]
96. k1_xboole_0 != sK1 <- (~1) [avatar component clause 94] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
97. ~1 [avatar split clause 64,94]
99. 2 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
101. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (2) [avatar component clause 99] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
102. 2 [avatar split clause 63,99]
160. v1_xboole_0(sK1) | ~m1_subset_1(sK3(sK1),sK0) [resolution 76,65] {a:1,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
163. 13 <=> m1_subset_1(sK3(sK1),sK0) [avatar definition]
165. ~m1_subset_1(sK3(sK1),sK0) <- (~13) [avatar component clause 163] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
167. 14 <=> v1_xboole_0(sK1) [avatar definition]
169. v1_xboole_0(sK1) <- (14) [avatar component clause 167] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
170. ~13 | 14 [avatar split clause 160,167,163]
202. ~r2_hidden(X1,X0) | m1_subset_1(X1,X0) [subsumption resolution 78,75] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
213. ~r2_hidden(X3,sK1) | r2_hidden(X3,sK0) <- (2) [resolution 81,101] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
217. r2_hidden(sK3(sK1),sK0) | v1_xboole_0(sK1) <- (2) [resolution 213,76] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
219. 19 <=> r2_hidden(sK3(sK1),sK0) [avatar definition]
221. r2_hidden(sK3(sK1),sK0) <- (19) [avatar component clause 219] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
222. 14 | 19 | ~2 [avatar split clause 217,99,219,167]
224. m1_subset_1(sK3(sK1),sK0) <- (19) [resolution 221,202] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
228. $false <- (~13, 19) [subsumption resolution 224,165] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
229. 13 | ~19 [avatar contradiction clause 228]
230. $false <- (~1, 14) [unit resulting resolution 96,169,74] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
234. 1 | ~14 [avatar contradiction clause 230]
237. $false [avatar sat refutation 97,102,170,222,229,234] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:108,thDist:-108}
% SZS output end Proof for subset_1__t4_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% Introduced skolems: 10
% 
% >>> Saturation
% Initial clauses: 30
% Generated clauses: 94
% Active clauses: 40
% Passive clauses: 52
% Final active clauses: 38
% Final passive clauses: 9
% Main loop iterations started: 43
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 32
% Unit resulting resolution: 4
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 10
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

