i: [1,80,0,3,1,1,0,0] t33_subset_1
a: [3,133,0,3,1,80]
s: 133
i: [1,79,0,3,1,1,0,0] t33_subset_1
a: [3,138,0,3,1,79]
s: 138
i: [1,81,0,5,1,1,0,0] t33_subset_1
a: [3,128,0,5,1,81]
s: 128
i: [1,82,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,143,0,2,1,82]
s: 143
i: [1,83,0,2,1,0,0,1] fc1_xboole_0
a: [3,148,0,2,1,83]
s: 148
i: [1,115,0,2,1,0,0,2] rc2_zfmisc_1
a: [3,158,0,2,1,115]
s: 158
i: [1,116,0,2,1,0,0,2] rc2_xboole_0
a: [3,163,0,2,1,116]
s: 163
i: [1,117,0,2,1,0,0,2] rc1_xboole_0
a: [3,168,0,2,1,117]
s: 168
i: [1,119,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,173,0,2,1,119]
s: 173
i: [1,118,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,178,0,2,1,118]
s: 178
i: [1,85,0,3,1,0,0,1] fc2_xboole_0
s: 85
i: [1,86,0,3,1,0,0,1] fc1_subset_1
s: 86
i: [1,87,0,3,1,0,0,1] fc2_zfmisc_1
s: 87
i: [1,112,0,3,1,0,0,1] rc2_subset_1
s: 112
i: [1,113,0,3,1,0,0,255] reflexivity_r1_tarski
s: 113
i: [1,84,0,3,1,0,0,255] d2_xboole_0
a: [3,153,0,3,1,84]
s: 153
r: 143
i: [1,88,0,4,2,0,0,2] cc2_zfmisc_1
s: 88
i: [1,91,0,4,2,0,0,2] cc1_zfmisc_1
s: 91
d: [2,181,1,2,1,1,40,88,148]
a: [3,187,1,2,1,181]
s: 187
i: [1,110,0,4,1,0,0,1] existence_m1_subset_1
s: 110
i: [1,121,0,4,1,0,0,1] d1_tarski
s: 121
d: [2,183,1,2,1,1,40,88,168]
a: [3,192,1,2,1,183]
s: 192
i: [1,90,0,5,2,0,0,1] rc1_subset_1
s: 90
d: [2,182,1,3,1,0,40,88,112]
s: 182
i: [1,92,0,5,2,0,0,1] t6_boole
s: 92
i: [1,109,0,5,2,0,0,2] t7_boole
s: 109
i: [1,111,0,5,1,0,0,1] rc2_subset_1
s: 111
i: [1,114,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 114
d: [2,198,1,4,1,0,40,114,121]
s: 198
i: [1,89,0,7,2,0,0,1] rc1_subset_1
s: 89
i: [1,95,0,7,3,0,0,1] d1_subset_1
s: 95
d: [2,199,1,4,2,1,40,95,138]
a: [3,206,1,2,1,199]
s: 206
i: [1,96,0,7,3,0,0,1] d1_subset_1
s: 96
d: [2,201,1,5,2,0,40,95,110]
s: 201
i: [1,108,0,7,3,0,0,1] t8_boole
s: 108
i: [1,122,0,7,2,0,0,1] d1_tarski
s: 122
d: [2,215,2,5,2,0,40,201,88]
s: 215
i: [1,123,0,7,2,0,0,255] d1_zfmisc_1
s: 123
d: [2,219,2,5,1,0,40,123,198]
s: 219
i: [1,124,0,7,2,0,0,255] d1_zfmisc_1
s: 124
i: [1,93,0,8,3,0,0,1] d1_subset_1
s: 93
d: [2,222,1,5,2,1,40,93,138]
d: [2,226,1,3,1,2,21,222,206]
a: [3,230,1,3,1,226]
s: 230
i: [1,94,0,8,3,0,0,1] d1_subset_1
s: 94
d: [2,234,2,3,1,1,40,230,114]
a: [3,239,2,3,1,234]
s: 239
d: [2,241,0,6,2,0,21,94,109]
s: 241
i: [1,98,0,8,2,0,0,255] d3_tarski
s: 98
d: [2,242,1,4,1,0,40,241,121]
s: 242
i: [1,99,0,8,2,0,0,255] d3_tarski
s: 99
d: [2,225,1,8,2,0,40,93,111]
d: [2,233,1,5,1,0,21,225,86]
s: 233
i: [1,97,0,9,3,0,0,255] d3_tarski
s: 97
i: [1,102,0,14,3,0,0,1] d1_tarski
s: 102
d: [2,255,2,4,1,0,40,233,124]
s: 255
i: [1,103,0,14,3,0,0,1] d1_tarski
s: 103
d: [2,247,1,5,2,0,40,98,109]
s: 247
i: [1,106,0,14,3,0,0,255] d1_zfmisc_1
s: 106
i: [1,107,0,14,3,0,0,255] d1_zfmisc_1
s: 107
d: [2,257,2,5,1,0,40,233,114]
s: 257
d: [2,224,1,6,2,0,40,93,110]
s: 224
d: [2,275,2,8,2,0,40,224,122]
d: [2,277,2,5,1,0,21,275,85]
s: 277
d: [2,194,1,7,3,0,44,92,92]
s: 194
d: [2,217,1,7,2,0,40,123,114]
s: 217
d: [2,276,2,8,2,0,40,224,124]
d: [2,278,2,5,1,0,21,276,86]
s: 278
d: [2,223,1,10,3,0,40,93,89]
d: [2,232,1,7,2,0,21,223,86]
s: 232
d: [2,273,2,6,2,0,40,224,114]
s: 273
d: [2,243,1,7,2,0,40,241,123]
s: 243
d: [2,245,1,8,2,0,40,98,241]
s: 245
d: [2,291,2,4,1,1,40,243,128]
a: [3,297,2,4,1,291]
s: 297
d: [2,246,1,8,2,0,40,98,114]
s: 246
d: [2,286,2,6,2,0,40,232,124]
s: 286
d: [2,248,1,10,2,0,40,98,122]
s: 248
d: [2,249,1,10,2,0,40,98,124]
s: 249
d: [2,288,2,7,2,0,40,232,114]
s: 288
d: [2,253,1,10,2,0,40,99,123]
s: 253
d: [2,305,2,11,3,0,44,99,248]
d: [2,308,2,7,2,0,18,305]
s: 308
d: [2,317,3,0,0,2,52,230,297,308]
e: 317
d: [2,319,3,3,1,1,40,308,297]
d: [2,322,3,0,0,2,21,319,230]
e: 322
r: 206
r: 230
a: [3,205,1,2,1,199]
s: 205
d: [2,326,2,0,0,2,52,133,205,92]
e: 326
d: [2,327,2,0,0,3,52,148,133,205,108]
e: 327
d: [2,328,2,0,0,3,52,148,133,205,108]
e: 328
f: 328,319,222,291,83,79,80,81
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t33_subset_1
% SZS output start Proof for subset_1__t33_subset_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,X0) => (k1_xboole_0 != X0 => m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X0)))) [input t33_subset_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,X0) => (k1_xboole_0 != X0 => m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X0)))) [negated conjecture 1]
6. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
7. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
8. ! [X0] : ! [X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [input d1_zfmisc_1]
10. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
18. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
29. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
30. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [input t8_boole]
31. ~! [X0,X1] : (m1_subset_1(X1,X0) => (k1_xboole_0 != X0 => m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X0)))) [flattening 2]
32. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 6]
33. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 10]
34. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 7]
35. ! [X0,X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [flattening 8]
36. ! [X0,X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [flattening 30]
37. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 29]
39. ? [X0,X1] : ((~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X0)) & k1_xboole_0 != X0) & m1_subset_1(X1,X0)) [ennf transformation 31]
40. ? [X0,X1] : (~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X0)) & k1_xboole_0 != X0 & m1_subset_1(X1,X0)) [flattening 39]
45. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 32]
46. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 33]
47. ! [X0,X1] : (~v1_xboole_0(X1) | X0 = X1 | ~v1_xboole_0(X0)) [ennf transformation 36]
48. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 37]
50. ? [X0,X1] : (~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X0)) & k1_xboole_0 != X0 & m1_subset_1(X1,X0)) => (~m1_subset_1(k1_tarski(sK1),k1_zfmisc_1(sK0)) & k1_xboole_0 != sK0 & m1_subset_1(sK1,sK0)) [choice axiom]
51. ~m1_subset_1(k1_tarski(sK1),k1_zfmisc_1(sK0)) & k1_xboole_0 != sK0 & m1_subset_1(sK1,sK0) [skolemisation 40,50]
54. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 45]
55. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 46]
56. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 55]
57. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK3(X0,X1),X1) & r2_hidden(sK3(X0,X1),X0))) [choice axiom]
58. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK3(X0,X1),X1) & r2_hidden(sK3(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 56,57]
59. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 34]
60. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 59]
61. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
62. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 60,61]
63. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ~r1_tarski(X2,X0)) & (r1_tarski(X2,X0) | ~r2_hidden(X2,X1))) | k1_zfmisc_1(X0) != X1)) [nnf transformation 35]
64. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [rectify 63]
65. ! [X1,X0] : (? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1))) => ((~r1_tarski(sK5(X0,X1),X0) | ~r2_hidden(sK5(X0,X1),X1)) & (r1_tarski(sK5(X0,X1),X0) | r2_hidden(sK5(X0,X1),X1)))) [choice axiom]
66. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ((~r1_tarski(sK5(X0,X1),X0) | ~r2_hidden(sK5(X0,X1),X1)) & (r1_tarski(sK5(X0,X1),X0) | r2_hidden(sK5(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [skolemisation 64,65]
79. m1_subset_1(sK1,sK0) [cnf transformation 51] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
80. k1_xboole_0 != sK0 [cnf transformation 51] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
81. ~m1_subset_1(k1_tarski(sK1),k1_zfmisc_1(sK0)) [cnf transformation 51] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
83. v1_xboole_0(k1_xboole_0) [cnf transformation 18] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
93. ~m1_subset_1(X1,X0) | r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 54] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
94. m1_subset_1(X1,X0) | ~r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 54] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
98. r2_hidden(sK3(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 58] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
99. ~r2_hidden(sK3(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 58] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
100. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 62] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
105. r2_hidden(X3,X1) | ~r1_tarski(X3,X0) | k1_zfmisc_1(X0) != X1 [cnf transformation 66] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
108. X0 = X1 | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 47] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
109. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 48] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
122. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 100] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
123. r2_hidden(X3,k1_zfmisc_1(X0)) | ~r1_tarski(X3,X0) [equality resolution 105] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
126. 1 <=> m1_subset_1(k1_tarski(sK1),k1_zfmisc_1(sK0)) [avatar definition]
128. ~m1_subset_1(k1_tarski(sK1),k1_zfmisc_1(sK0)) <- (~1) [avatar component clause 126] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
129. ~1 [avatar split clause 81,126]
131. 2 <=> k1_xboole_0 = sK0 [avatar definition]
133. k1_xboole_0 != sK0 <- (~2) [avatar component clause 131] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
134. ~2 [avatar split clause 80,131]
136. 3 <=> m1_subset_1(sK1,sK0) [avatar definition]
138. m1_subset_1(sK1,sK0) <- (3) [avatar component clause 136] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
139. 3 [avatar split clause 79,136]
146. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
148. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 146] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
149. 5 [avatar split clause 83,146]
204. 14 <=> v1_xboole_0(sK0) [avatar definition]
205. v1_xboole_0(sK0) <- (14) [avatar component clause 204] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
222. r2_hidden(sK1,sK0) | v1_xboole_0(sK0) <- (3) [resolution 93,138] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
228. 16 <=> r2_hidden(sK1,sK0) [avatar definition]
241. ~r2_hidden(X1,X0) | m1_subset_1(X1,X0) [subsumption resolution 94,109] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
243. m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r1_tarski(X1,X2) [resolution 241,123] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
248. sK3(k1_tarski(X6),X7) = X6 | r1_tarski(k1_tarski(X6),X7) [resolution 98,122] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
291. ~r1_tarski(k1_tarski(sK1),sK0) <- (~1) [resolution 243,128] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
295. 18 <=> r1_tarski(k1_tarski(sK1),sK0) [avatar definition]
297. ~r1_tarski(k1_tarski(sK1),sK0) <- (~18) [avatar component clause 295] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
298. ~18 | 1 [avatar split clause 291,126,295]
305. ~r2_hidden(X4,X5) | r1_tarski(k1_tarski(X4),X5) | r1_tarski(k1_tarski(X4),X5) [superposition 99,248] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:6,thDist:-6}
308. r1_tarski(k1_tarski(X4),X5) | ~r2_hidden(X4,X5) [duplicate literal removal 305] {a:2,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
319. ~r2_hidden(sK1,sK0) <- (~18) [resolution 308,297] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
324. 14 | 16 | ~3 [avatar split clause 222,136,228,204]
325. ~16 | 18 [avatar split clause 319,295,228]
328. $false <- (~2, 5, 14) [unit resulting resolution 148,133,205,108] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
332. 2 | ~5 | ~14 [avatar contradiction clause 328]
333. $false [avatar sat refutation 129,134,139,149,298,324,325,332] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:128,thDist:-128}
% SZS output end Proof for subset_1__t33_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% Introduced skolems: 12
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 177
% Active clauses: 72
% Passive clauses: 93
% Final active clauses: 69
% Final passive clauses: 18
% Main loop iterations started: 75
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 7
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Equational tautologies: 1
% Forward subsumptions: 36
% 
% >>> Generating Inferences
% Binary resolution: 78
% Unit resulting resolution: 4
% Forward superposition: 2
% Backward superposition: 10
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 11
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

