i: [1,88,0,2,1,1,0,0] t18_boolealg
a: [3,158,0,2,1,88]
s: 158
i: [1,87,0,2,1,1,0,0] t18_boolealg
a: [3,163,0,2,1,87]
s: 163
i: [1,86,0,2,1,1,0,0] t18_boolealg
a: [3,168,0,2,1,86]
s: 168
i: [1,85,0,2,1,1,0,0] t18_boolealg
a: [3,173,0,2,1,85]
s: 173
i: [1,92,0,8,2,1,0,0] t18_boolealg
a: [3,134,0,4,1,92]
s: 134
i: [1,91,0,4,1,1,0,0] t18_boolealg
a: [3,143,0,4,1,91]
s: 143
i: [1,90,0,4,1,1,0,0] t18_boolealg
a: [3,148,0,4,1,90]
s: 148
i: [1,89,0,4,1,1,0,0] t18_boolealg
a: [3,153,0,4,1,89]
s: 153
i: [1,93,0,7,1,1,0,0] t18_boolealg
a: [3,129,0,7,1,93]
s: 129
i: [1,124,0,2,1,0,0,2] existence_l1_lattices
a: [3,178,0,2,1,124]
s: 178
i: [1,125,0,2,1,0,0,1] existence_l3_lattices
a: [3,183,0,2,1,125]
s: 183
i: [1,94,0,4,2,0,0,255] dt_l3_lattices
s: 94
i: [1,106,0,4,1,0,0,1] existence_m1_subset_1
s: 106
d: [2,185,1,2,1,1,40,94,158]
a: [3,190,1,2,1,185]
s: 190
i: [1,98,0,8,4,0,0,255] cc1_lattices
s: 98
i: [1,99,0,8,4,0,0,255] cc1_lattices
s: 99
d: [2,197,1,6,3,1,40,98,158]
d: [2,199,1,4,2,2,21,197,168]
d: [2,200,1,2,1,3,21,199,173]
a: [3,204,1,2,1,200]
s: 204
i: [1,100,0,8,4,0,0,255] cc1_lattices
s: 100
d: [2,219,1,6,3,1,40,99,158]
d: [2,221,1,4,2,2,21,219,168]
d: [2,222,1,2,1,3,21,221,173]
a: [3,226,1,2,1,222]
s: 226
i: [1,101,0,8,4,0,0,255] cc1_lattices
s: 101
i: [1,102,0,8,4,0,0,255] cc1_lattices
s: 102
d: [2,233,1,6,3,1,40,100,158]
d: [2,235,1,4,2,2,21,233,168]
d: [2,236,1,2,1,3,21,235,173]
a: [3,240,1,2,1,236]
s: 240
i: [1,103,0,8,4,0,0,255] cc1_lattices
s: 103
d: [2,247,1,6,3,1,40,101,158]
d: [2,249,1,4,2,2,21,247,168]
d: [2,250,1,2,1,3,21,249,173]
a: [3,254,1,2,1,250]
s: 254
i: [1,96,0,18,9,0,0,255] cc2_lattices
s: 96
i: [1,111,0,19,5,0,0,255] dt_k2_lattices
s: 111
d: [2,261,1,6,3,1,40,102,158]
d: [2,263,1,4,2,2,21,261,168]
d: [2,264,1,2,1,3,21,263,173]
a: [3,268,1,2,1,264]
s: 268
i: [1,107,0,20,7,0,0,255] rd2_lattices
s: 107
d: [2,275,1,6,3,1,40,103,158]
d: [2,277,1,4,2,2,21,275,168]
d: [2,278,1,2,1,3,21,277,173]
a: [3,282,1,2,1,278]
s: 282
i: [1,108,0,21,6,0,0,1] dt_k4_lattices
s: 108
i: [1,112,0,22,7,0,0,1] symmetry_r4_boolealg
s: 112
d: [2,292,1,24,6,1,40,112,129]
d: [2,293,1,20,5,2,21,292,153]
d: [2,294,1,18,4,3,21,293,158]
d: [2,295,1,16,3,4,21,294,168]
d: [2,296,1,14,2,5,21,295,173]
a: [3,300,1,7,1,296]
s: 300
d: [2,306,2,0,0,6,52,173,240,190,148,143,300,108]
e: 306
d: [2,307,2,14,5,1,40,300,108]
d: [2,309,2,10,4,2,21,307,143]
d: [2,310,2,6,3,3,21,309,148]
d: [2,311,2,4,2,4,21,310,190]
d: [2,312,2,2,1,5,21,311,240]
d: [2,313,2,0,0,6,21,312,173]
e: 313
r: 300
i: [1,113,0,22,7,0,0,1] symmetry_r4_boolealg
s: 113
a: [3,299,1,7,1,296]
s: 299
i: [1,114,0,22,7,0,0,2] symmetry_r3_boolealg
s: 114
i: [1,115,0,22,7,0,0,2] symmetry_r3_boolealg
s: 115
a: [3,304,1,7,1,296]
s: 304
i: [1,116,0,22,7,0,0,255] redefinition_r4_boolealg
s: 116
d: [2,316,1,24,6,1,40,116,129]
d: [2,318,1,17,5,2,21,316,299]
d: [2,319,1,13,4,3,21,318,153]
d: [2,320,1,11,3,4,21,319,158]
d: [2,321,1,9,2,5,21,320,168]
d: [2,322,1,7,1,6,21,321,173]
a: [3,326,1,7,1,322]
s: 326
i: [1,117,0,22,7,0,0,255] redefinition_r4_boolealg
s: 117
i: [1,118,0,22,7,0,0,255] redefinition_r4_boolealg
s: 118
d: [2,317,2,24,6,1,40,116,304]
d: [2,328,2,20,5,2,21,317,153]
d: [2,329,2,13,4,3,21,328,299]
d: [2,330,2,11,3,4,21,329,158]
d: [2,331,2,9,2,5,21,330,168]
d: [2,332,2,7,1,6,21,331,173]
a: [3,336,2,7,1,332]
s: 336
i: [1,119,0,22,7,0,0,255] redefinition_r4_boolealg
s: 119
d: [2,186,1,2,1,1,40,94,183]
a: [3,195,1,2,1,186]
s: 195
i: [1,120,0,22,7,0,0,255] redefinition_r3_boolealg
s: 120
i: [1,121,0,22,7,0,0,255] redefinition_r3_boolealg
s: 121
d: [2,340,2,24,6,1,40,121,326]
d: [2,342,2,17,5,2,21,340,299]
d: [2,343,2,13,4,3,21,342,153]
d: [2,344,2,11,3,4,21,343,158]
d: [2,345,2,9,2,5,21,344,168]
d: [2,346,2,7,1,6,21,345,173]
a: [3,350,2,7,1,346]
s: 350
i: [1,122,0,22,7,0,0,255] redefinition_r3_boolealg
s: 122
d: [2,341,3,24,6,1,40,121,336]
d: [2,352,3,20,5,2,21,341,153]
d: [2,353,3,13,4,3,21,352,299]
d: [2,354,3,11,3,4,21,353,158]
d: [2,355,3,9,2,5,21,354,168]
d: [2,356,3,7,1,6,21,355,173]
a: [3,360,3,7,1,356]
s: 360
i: [1,123,0,22,7,0,0,255] redefinition_r3_boolealg
s: 123
i: [1,109,0,23,6,0,0,255] redefinition_k4_lattices
s: 109
d: [2,198,1,6,3,1,40,98,183]
a: [3,217,1,2,1,198]
s: 217
i: [1,110,0,23,6,0,0,255] commutativity_k4_lattices
s: 110
i: [1,104,0,31,9,0,0,1] t14_boolealg
s: 104
i: [1,105,0,31,9,0,0,1] t14_boolealg
s: 105
d: [2,380,1,45,14,0,44,104,110]
d: [2,383,1,35,11,0,18,380]
d: [2,397,1,33,10,0,21,383,100]
d: [2,398,1,31,9,0,21,397,94]
s: 398
d: [2,378,3,24,8,1,40,104,350]
d: [2,385,3,20,7,2,21,378,143]
d: [2,386,3,16,6,3,21,385,148]
d: [2,387,3,12,5,4,21,386,153]
d: [2,388,3,10,4,5,21,387,158]
d: [2,389,3,8,3,6,21,388,163]
d: [2,390,3,6,2,7,21,389,168]
d: [2,391,3,4,1,8,21,390,173]
a: [3,395,3,4,1,391]
s: 395
d: [2,381,1,45,14,0,44,104,110]
d: [2,382,1,35,11,0,18,381]
d: [2,399,1,33,10,0,21,382,100]
d: [2,400,1,31,9,0,21,399,94]
s: 400
d: [2,401,3,24,8,1,40,105,350]
d: [2,408,3,20,7,2,21,401,143]
d: [2,409,3,16,6,3,21,408,148]
d: [2,410,3,12,5,4,21,409,153]
d: [2,411,3,10,4,5,21,410,158]
d: [2,412,3,8,3,6,21,411,163]
d: [2,413,3,6,2,7,21,412,168]
d: [2,414,3,4,1,8,21,413,173]
a: [3,418,3,4,1,414]
s: 418
d: [2,420,4,18,6,1,40,395,120]
d: [2,422,4,14,5,2,21,420,148]
d: [2,423,4,10,4,3,21,422,153]
d: [2,424,4,8,3,4,21,423,158]
d: [2,425,4,6,2,5,21,424,168]
d: [2,426,4,4,1,6,21,425,173]
a: [3,430,4,4,1,426]
s: 430
d: [2,464,5,0,0,7,52,173,168,158,153,148,134,430,117]
e: 464
d: [2,466,5,18,6,1,40,430,117]
d: [2,468,5,14,5,2,21,466,134]
d: [2,469,5,10,4,3,21,468,148]
d: [2,470,5,6,3,4,21,469,153]
d: [2,471,5,4,2,5,21,470,158]
d: [2,472,5,2,1,6,21,471,168]
d: [2,473,5,0,0,7,21,472,173]
e: 473
r: 134
a: [3,133,0,4,1,92]
s: 133
a: [3,138,0,4,1,92]
s: 138
d: [2,482,1,18,6,1,40,133,112]
d: [2,483,1,14,5,2,21,482,148]
d: [2,484,1,10,4,3,21,483,153]
d: [2,485,1,8,3,4,21,484,158]
d: [2,486,1,6,2,5,21,485,168]
d: [2,487,1,4,1,6,21,486,173]
a: [3,491,1,4,1,487]
s: 491
d: [2,493,2,18,6,1,40,491,116]
d: [2,495,2,14,5,2,21,493,153]
d: [2,496,2,10,4,3,21,495,148]
d: [2,497,2,8,3,4,21,496,158]
d: [2,498,2,6,2,5,21,497,168]
d: [2,499,2,4,1,6,21,498,173]
a: [3,503,2,4,1,499]
s: 503
d: [2,421,4,18,6,1,40,395,114]
d: [2,432,4,14,5,2,21,421,148]
d: [2,433,4,10,4,3,21,432,153]
d: [2,434,4,8,3,4,21,433,158]
d: [2,435,4,6,2,5,21,434,168]
d: [2,436,4,4,1,6,21,435,173]
a: [3,440,4,4,1,436]
s: 440
d: [2,442,4,18,6,1,40,418,120]
d: [2,444,4,14,5,2,21,442,143]
d: [2,445,4,10,4,3,21,444,153]
d: [2,446,4,8,3,4,21,445,158]
d: [2,447,4,6,2,5,21,446,168]
d: [2,448,4,4,1,6,21,447,173]
a: [3,452,4,4,1,448]
s: 452
d: [2,515,5,0,0,7,52,173,168,158,153,143,138,452,117]
e: 515
d: [2,517,5,18,6,1,40,452,117]
d: [2,519,5,14,5,2,21,517,138]
d: [2,520,5,10,4,3,21,519,143]
d: [2,521,5,6,3,4,21,520,153]
d: [2,522,5,4,2,5,21,521,158]
d: [2,523,5,2,1,6,21,522,168]
d: [2,524,5,0,0,7,21,523,173]
e: 524
d: [2,475,5,14,5,2,21,466,148]
d: [2,476,5,10,4,3,21,475,153]
d: [2,477,5,8,3,4,21,476,158]
d: [2,478,5,6,2,5,21,477,168]
d: [2,479,5,4,1,6,21,478,173]
f: 524,479,448,426,414,391,346,322,313,236,185,85,86,87,88,89,90,91,92,93
% Refutation found. Thanks to Tanya!
% SZS status Theorem for boolealg__t18_boolealg
% SZS output start Proof for boolealg__t18_boolealg
1. ! [X0] : ((l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ~(r4_boolealg(X0,X1,k4_lattices(X0,X2,X3)) & ~(r4_boolealg(X0,X1,X3) & r4_boolealg(X0,X1,X2))))))) [input t18_boolealg]
2. ~! [X0] : ((l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ~(r4_boolealg(X0,X1,k4_lattices(X0,X2,X3)) & ~(r4_boolealg(X0,X1,X3) & r4_boolealg(X0,X1,X2))))))) [negated conjecture 1]
21. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (r3_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2))) [input redefinition_r3_boolealg]
22. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (r4_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2))) [input redefinition_r4_boolealg]
24. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_lattices(X0) & v6_lattices(X0) & ~v2_struct_0(X0)) => m1_subset_1(k4_lattices(X0,X1,X2),u1_struct_0(X0))) [input dt_k4_lattices]
25. ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))) [input dt_l3_lattices]
28. ! [X0] : (l3_lattices(X0) => ((v10_lattices(X0) & ~v2_struct_0(X0)) => (v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)))) [input cc1_lattices]
29. ! [X0] : ((l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r3_boolealg(X0,X1,k4_lattices(X0,X2,X3)) => (r3_boolealg(X0,X1,X3) & r3_boolealg(X0,X1,X2))))))) [input t14_boolealg]
33. ! [X0] : (l3_lattices(X0) => l1_lattices(X0)) [pure predicate removal 25]
35. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((r4_boolealg(X0,X1,k4_lattices(X0,X2,X3)) & (~r4_boolealg(X0,X1,X3) | ~r4_boolealg(X0,X1,X2))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
36. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (r4_boolealg(X0,X1,k4_lattices(X0,X2,X3)) & (~r4_boolealg(X0,X1,X3) | ~r4_boolealg(X0,X1,X2)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) [flattening 35]
37. ! [X0] : (l1_lattices(X0) | ~l3_lattices(X0)) [ennf transformation 33]
40. ! [X0] : (((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)) | (~v10_lattices(X0) | v2_struct_0(X0))) | ~l3_lattices(X0)) [ennf transformation 28]
41. ! [X0] : ((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)) | ~v10_lattices(X0) | v2_struct_0(X0) | ~l3_lattices(X0)) [flattening 40]
42. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r3_boolealg(X0,X1,X3) & r3_boolealg(X0,X1,X2)) | ~r3_boolealg(X0,X1,k4_lattices(X0,X2,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 29]
43. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r3_boolealg(X0,X1,X3) & r3_boolealg(X0,X1,X2)) | ~r3_boolealg(X0,X1,k4_lattices(X0,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 42]
46. ! [X0,X1,X2] : (m1_subset_1(k4_lattices(X0,X1,X2),u1_struct_0(X0)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_lattices(X0) | ~v6_lattices(X0) | v2_struct_0(X0))) [ennf transformation 24]
47. ! [X0,X1,X2] : (m1_subset_1(k4_lattices(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_lattices(X0) | ~v6_lattices(X0) | v2_struct_0(X0)) [flattening 46]
62. ! [X0,X1,X2] : ((r4_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 22]
63. ! [X0,X1,X2] : ((r4_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 62]
66. ! [X0,X1,X2] : ((r3_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 21]
67. ! [X0,X1,X2] : ((r3_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 66]
70. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (r4_boolealg(X0,X1,k4_lattices(X0,X2,X3)) & (~r4_boolealg(X0,X1,X3) | ~r4_boolealg(X0,X1,X2)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (r4_boolealg(sK0,X1,k4_lattices(sK0,X2,X3)) & (~r4_boolealg(sK0,X1,X3) | ~r4_boolealg(sK0,X1,X2)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l3_lattices(sK0) & v13_lattices(sK0) & v10_lattices(sK0) & ~v2_struct_0(sK0)) [choice axiom]
71. ? [X1] : (? [X2] : (? [X3] : (r4_boolealg(sK0,X1,k4_lattices(sK0,X2,X3)) & (~r4_boolealg(sK0,X1,X3) | ~r4_boolealg(sK0,X1,X2)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (r4_boolealg(sK0,sK1,k4_lattices(sK0,X2,X3)) & (~r4_boolealg(sK0,sK1,X3) | ~r4_boolealg(sK0,sK1,X2)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
72. ? [X2] : (? [X3] : (r4_boolealg(sK0,sK1,k4_lattices(sK0,X2,X3)) & (~r4_boolealg(sK0,sK1,X3) | ~r4_boolealg(sK0,sK1,X2)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (r4_boolealg(sK0,sK1,k4_lattices(sK0,sK2,X3)) & (~r4_boolealg(sK0,sK1,X3) | ~r4_boolealg(sK0,sK1,sK2)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
73. ? [X3] : (r4_boolealg(sK0,sK1,k4_lattices(sK0,sK2,X3)) & (~r4_boolealg(sK0,sK1,X3) | ~r4_boolealg(sK0,sK1,sK2)) & m1_subset_1(X3,u1_struct_0(sK0))) => (r4_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) & (~r4_boolealg(sK0,sK1,sK3) | ~r4_boolealg(sK0,sK1,sK2)) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
74. (((r4_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) & (~r4_boolealg(sK0,sK1,sK3) | ~r4_boolealg(sK0,sK1,sK2)) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l3_lattices(sK0) & v13_lattices(sK0) & v10_lattices(sK0) & ~v2_struct_0(sK0) [skolemisation 36,73,72,71,70]
77. ! [X0,X1,X2] : (((r4_boolealg(X0,X1,X2) | ~r2_boolealg(X0,X1,X2)) & (r2_boolealg(X0,X1,X2) | ~r4_boolealg(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [nnf transformation 63]
79. ! [X0,X1,X2] : (((r3_boolealg(X0,X1,X2) | ~r2_boolealg(X0,X1,X2)) & (r2_boolealg(X0,X1,X2) | ~r3_boolealg(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [nnf transformation 67]
85. ~v2_struct_0(sK0) [cnf transformation 74] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
86. v10_lattices(sK0) [cnf transformation 74] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
87. v13_lattices(sK0) [cnf transformation 74] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
88. l3_lattices(sK0) [cnf transformation 74] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
89. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 74] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
90. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 74] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
91. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 74] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
92. ~r4_boolealg(sK0,sK1,sK3) | ~r4_boolealg(sK0,sK1,sK2) [cnf transformation 74] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
93. r4_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) [cnf transformation 74] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
94. ~l3_lattices(X0) | l1_lattices(X0) [cnf transformation 37] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
100. ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | v6_lattices(X0) [cnf transformation 41] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
104. ~r3_boolealg(X0,X1,k4_lattices(X0,X2,X3)) | r3_boolealg(X0,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 43] {a:0,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
105. ~r3_boolealg(X0,X1,k4_lattices(X0,X2,X3)) | r3_boolealg(X0,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 43] {a:0,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
108. m1_subset_1(k4_lattices(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_lattices(X0) | ~v6_lattices(X0) | v2_struct_0(X0) [cnf transformation 47] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
116. ~r4_boolealg(X0,X1,X2) | r2_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 77] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
117. ~r2_boolealg(X0,X1,X2) | r4_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 77] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
120. ~r3_boolealg(X0,X1,X2) | r2_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 79] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
121. ~r2_boolealg(X0,X1,X2) | r3_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 79] {a:0,w:22,wCS:8800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
127. 1 <=> r4_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) [avatar definition]
129. r4_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) <- (1) [avatar component clause 127] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
130. 1 [avatar split clause 93,127]
132. 2 <=> r4_boolealg(sK0,sK1,sK2) [avatar definition]
136. 3 <=> r4_boolealg(sK0,sK1,sK3) [avatar definition]
138. ~r4_boolealg(sK0,sK1,sK3) <- (~3) [avatar component clause 136] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
139. ~2 | ~3 [avatar split clause 92,136,132]
141. 4 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
143. m1_subset_1(sK3,u1_struct_0(sK0)) <- (4) [avatar component clause 141] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
144. 4 [avatar split clause 91,141]
146. 5 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
148. m1_subset_1(sK2,u1_struct_0(sK0)) <- (5) [avatar component clause 146] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
149. 5 [avatar split clause 90,146]
151. 6 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
153. m1_subset_1(sK1,u1_struct_0(sK0)) <- (6) [avatar component clause 151] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
154. 6 [avatar split clause 89,151]
156. 7 <=> l3_lattices(sK0) [avatar definition]
158. l3_lattices(sK0) <- (7) [avatar component clause 156] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
159. 7 [avatar split clause 88,156]
161. 8 <=> v13_lattices(sK0) [avatar definition]
163. v13_lattices(sK0) <- (8) [avatar component clause 161] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
164. 8 [avatar split clause 87,161]
166. 9 <=> v10_lattices(sK0) [avatar definition]
168. v10_lattices(sK0) <- (9) [avatar component clause 166] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
169. 9 [avatar split clause 86,166]
171. 10 <=> v2_struct_0(sK0) [avatar definition]
173. ~v2_struct_0(sK0) <- (~10) [avatar component clause 171] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
174. ~10 [avatar split clause 85,171]
185. l1_lattices(sK0) <- (7) [resolution 94,158] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
188. 13 <=> l1_lattices(sK0) [avatar definition]
190. l1_lattices(sK0) <- (13) [avatar component clause 188] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
191. 13 | ~7 [avatar split clause 185,156,188]
233. ~v10_lattices(sK0) | v2_struct_0(sK0) | v6_lattices(sK0) <- (7) [resolution 100,158] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
235. v2_struct_0(sK0) | v6_lattices(sK0) <- (7, 9) [subsumption resolution 233,168] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
236. v6_lattices(sK0) <- (7, 9, ~10) [subsumption resolution 235,173] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
238. 21 <=> v6_lattices(sK0) [avatar definition]
240. v6_lattices(sK0) <- (21) [avatar component clause 238] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
241. 21 | ~7 | ~9 | 10 [avatar split clause 236,171,166,156,238]
298. 29 <=> m1_subset_1(k4_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) [avatar definition]
299. m1_subset_1(k4_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) <- (29) [avatar component clause 298] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
300. ~m1_subset_1(k4_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) <- (~29) [avatar component clause 298] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
307. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_lattices(sK0) | ~v6_lattices(sK0) | v2_struct_0(sK0) <- (~29) [resolution 300,108] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
309. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_lattices(sK0) | ~v6_lattices(sK0) | v2_struct_0(sK0) <- (4, ~29) [subsumption resolution 307,143] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
310. ~l1_lattices(sK0) | ~v6_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, ~29) [subsumption resolution 309,148] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
311. ~v6_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 13, ~29) [subsumption resolution 310,190] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
312. v2_struct_0(sK0) <- (4, 5, 13, 21, ~29) [subsumption resolution 311,240] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
313. $false <- (4, 5, ~10, 13, 21, ~29) [subsumption resolution 312,173] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
314. ~4 | ~5 | 10 | ~13 | ~21 | 29 [avatar contradiction clause 313]
316. r2_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | ~m1_subset_1(k4_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1) [resolution 116,129] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
318. r2_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1, 29) [subsumption resolution 316,299] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
319. r2_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1, 6, 29) [subsumption resolution 318,153] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
320. r2_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1, 6, 7, 29) [subsumption resolution 319,158] {a:1,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
321. r2_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | v2_struct_0(sK0) <- (1, 6, 7, 9, 29) [subsumption resolution 320,168] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
322. r2_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) <- (1, 6, 7, 9, ~10, 29) [subsumption resolution 321,173] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
324. 31 <=> r2_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) [avatar definition]
326. r2_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) <- (31) [avatar component clause 324] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
327. 31 | ~1 | ~6 | ~7 | ~9 | 10 | ~29 [avatar split clause 322,298,171,166,156,151,127,324]
340. r3_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | ~m1_subset_1(k4_lattices(sK0,sK2,sK3),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (31) [resolution 121,326] {a:2,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
342. r3_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (29, 31) [subsumption resolution 340,299] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
343. r3_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (6, 29, 31) [subsumption resolution 342,153] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
344. r3_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (6, 7, 29, 31) [subsumption resolution 343,158] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
345. r3_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) | v2_struct_0(sK0) <- (6, 7, 9, 29, 31) [subsumption resolution 344,168] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
346. r3_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) <- (6, 7, 9, ~10, 29, 31) [subsumption resolution 345,173] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
348. 33 <=> r3_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) [avatar definition]
350. r3_boolealg(sK0,sK1,k4_lattices(sK0,sK2,sK3)) <- (33) [avatar component clause 348] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
351. 33 | ~6 | ~7 | ~9 | 10 | ~29 | ~31 [avatar split clause 346,324,298,171,166,156,151,348]
378. r3_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (33) [resolution 104,350] {a:3,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
385. r3_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 33) [subsumption resolution 378,143] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
386. r3_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 33) [subsumption resolution 385,148] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
387. r3_boolealg(sK0,sK1,sK2) | ~l3_lattices(sK0) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 6, 33) [subsumption resolution 386,153] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
388. r3_boolealg(sK0,sK1,sK2) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 6, 7, 33) [subsumption resolution 387,158] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
389. r3_boolealg(sK0,sK1,sK2) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 6, 7, 8, 33) [subsumption resolution 388,163] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
390. r3_boolealg(sK0,sK1,sK2) | v2_struct_0(sK0) <- (4, 5, 6, 7, 8, 9, 33) [subsumption resolution 389,168] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
391. r3_boolealg(sK0,sK1,sK2) <- (4, 5, 6, 7, 8, 9, ~10, 33) [subsumption resolution 390,173] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
393. 35 <=> r3_boolealg(sK0,sK1,sK2) [avatar definition]
395. r3_boolealg(sK0,sK1,sK2) <- (35) [avatar component clause 393] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
396. 35 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 10 | ~33 [avatar split clause 391,348,171,166,161,156,151,146,141,393]
401. r3_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (33) [resolution 105,350] {a:3,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
408. r3_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 33) [subsumption resolution 401,143] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
409. r3_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 33) [subsumption resolution 408,148] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
410. r3_boolealg(sK0,sK1,sK3) | ~l3_lattices(sK0) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 6, 33) [subsumption resolution 409,153] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
411. r3_boolealg(sK0,sK1,sK3) | ~v13_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 6, 7, 33) [subsumption resolution 410,158] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
412. r3_boolealg(sK0,sK1,sK3) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 5, 6, 7, 8, 33) [subsumption resolution 411,163] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
413. r3_boolealg(sK0,sK1,sK3) | v2_struct_0(sK0) <- (4, 5, 6, 7, 8, 9, 33) [subsumption resolution 412,168] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
414. r3_boolealg(sK0,sK1,sK3) <- (4, 5, 6, 7, 8, 9, ~10, 33) [subsumption resolution 413,173] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
416. 36 <=> r3_boolealg(sK0,sK1,sK3) [avatar definition]
418. r3_boolealg(sK0,sK1,sK3) <- (36) [avatar component clause 416] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
419. 36 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | 10 | ~33 [avatar split clause 414,348,171,166,161,156,151,146,141,416]
420. r2_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (35) [resolution 395,120] {a:4,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
422. r2_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 35) [subsumption resolution 420,148] {a:4,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
423. r2_boolealg(sK0,sK1,sK2) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 6, 35) [subsumption resolution 422,153] {a:4,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
424. r2_boolealg(sK0,sK1,sK2) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 6, 7, 35) [subsumption resolution 423,158] {a:4,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
425. r2_boolealg(sK0,sK1,sK2) | v2_struct_0(sK0) <- (5, 6, 7, 9, 35) [subsumption resolution 424,168] {a:4,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
426. r2_boolealg(sK0,sK1,sK2) <- (5, 6, 7, 9, ~10, 35) [subsumption resolution 425,173] {a:4,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
428. 37 <=> r2_boolealg(sK0,sK1,sK2) [avatar definition]
430. r2_boolealg(sK0,sK1,sK2) <- (37) [avatar component clause 428] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
431. 37 | ~5 | ~6 | ~7 | ~9 | 10 | ~35 [avatar split clause 426,393,171,166,156,151,146,428]
442. r2_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (36) [resolution 418,120] {a:4,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
444. r2_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 36) [subsumption resolution 442,143] {a:4,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
445. r2_boolealg(sK0,sK1,sK3) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 6, 36) [subsumption resolution 444,153] {a:4,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
446. r2_boolealg(sK0,sK1,sK3) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (4, 6, 7, 36) [subsumption resolution 445,158] {a:4,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
447. r2_boolealg(sK0,sK1,sK3) | v2_struct_0(sK0) <- (4, 6, 7, 9, 36) [subsumption resolution 446,168] {a:4,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
448. r2_boolealg(sK0,sK1,sK3) <- (4, 6, 7, 9, ~10, 36) [subsumption resolution 447,173] {a:4,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
450. 39 <=> r2_boolealg(sK0,sK1,sK3) [avatar definition]
452. r2_boolealg(sK0,sK1,sK3) <- (39) [avatar component clause 450] {a:4,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
453. 39 | ~4 | ~6 | ~7 | ~9 | 10 | ~36 [avatar split clause 448,416,171,166,156,151,141,450]
466. r4_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (37) [resolution 430,117] {a:5,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
475. r4_boolealg(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 37) [subsumption resolution 466,148] {a:5,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
476. r4_boolealg(sK0,sK1,sK2) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 6, 37) [subsumption resolution 475,153] {a:5,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
477. r4_boolealg(sK0,sK1,sK2) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (5, 6, 7, 37) [subsumption resolution 476,158] {a:5,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
478. r4_boolealg(sK0,sK1,sK2) | v2_struct_0(sK0) <- (5, 6, 7, 9, 37) [subsumption resolution 477,168] {a:5,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
479. r4_boolealg(sK0,sK1,sK2) <- (5, 6, 7, 9, ~10, 37) [subsumption resolution 478,173] {a:5,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
480. 2 | ~5 | ~6 | ~7 | ~9 | 10 | ~37 [avatar split clause 479,428,171,166,156,151,146,132]
517. r4_boolealg(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (39) [resolution 452,117] {a:5,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
519. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~3, 39) [subsumption resolution 517,138] {a:5,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
520. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~3, 4, 39) [subsumption resolution 519,143] {a:5,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
521. ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~3, 4, 6, 39) [subsumption resolution 520,153] {a:5,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
522. ~v10_lattices(sK0) | v2_struct_0(sK0) <- (~3, 4, 6, 7, 39) [subsumption resolution 521,158] {a:5,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
523. v2_struct_0(sK0) <- (~3, 4, 6, 7, 9, 39) [subsumption resolution 522,168] {a:5,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
524. $false <- (~3, 4, 6, 7, 9, ~10, 39) [subsumption resolution 523,173] {a:5,w:0,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
525. 3 | ~4 | ~6 | ~7 | ~9 | 10 | ~39 [avatar contradiction clause 524]
526. $false [avatar sat refutation 130,139,144,149,154,159,164,169,174,191,241,314,327,351,396,419,431,453,480,525] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:558,thDist:-558}
% SZS output end Proof for boolealg__t18_boolealg
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 28
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 264
% Active clauses: 58
% Passive clauses: 67
% Final active clauses: 56
% Final passive clauses: 1
% Main loop iterations started: 66
% 
% >>> Simplifying Inferences
% Duplicate literals: 38
% Fw subsumption resolutions: 116
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 27
% 
% >>> Generating Inferences
% Binary resolution: 46
% Unit resulting resolution: 3
% Forward superposition: 7
% Backward superposition: 3
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 22
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 49
% SAT solver unit clauses: 10
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9722
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.002 s
% other: 0.003 s ( own 0.001 s ) 

