i: [1,66,0,2,1,1,0,0] t9_boolealg
a: [3,115,0,2,1,66]
s: 115
i: [1,65,0,2,1,1,0,0] t9_boolealg
a: [3,120,0,2,1,65]
s: 120
i: [1,64,0,2,1,1,0,0] t9_boolealg
a: [3,125,0,2,1,64]
s: 125
i: [1,63,0,2,1,1,0,0] t9_boolealg
a: [3,130,0,2,1,63]
s: 130
i: [1,67,0,4,1,1,0,0] t9_boolealg
a: [3,110,0,4,1,67]
s: 110
i: [1,69,0,5,1,1,0,0] t9_boolealg
a: [3,100,0,5,1,69]
s: 100
i: [1,68,0,5,1,1,0,0] t9_boolealg
a: [3,105,0,5,1,68]
s: 105
i: [1,92,0,2,1,0,0,2] existence_l1_lattices
a: [3,135,0,2,1,92]
s: 135
i: [1,93,0,2,1,0,0,3] existence_l2_lattices
a: [3,140,0,2,1,93]
s: 140
i: [1,94,0,2,1,0,0,1] existence_l3_lattices
a: [3,145,0,2,1,94]
s: 145
i: [1,70,0,4,2,0,0,2] dt_l3_lattices
s: 70
d: [2,147,1,2,1,1,40,70,115]
a: [3,152,1,2,1,147]
s: 152
i: [1,71,0,4,2,0,0,2] dt_l3_lattices
s: 71
d: [2,159,1,2,1,1,40,71,115]
a: [3,164,1,2,1,159]
s: 164
i: [1,84,0,4,1,0,0,1] existence_m1_subset_1
s: 84
i: [1,75,0,8,4,0,0,255] cc1_lattices
s: 75
d: [2,171,1,6,3,1,40,75,115]
d: [2,173,1,4,2,2,21,171,125]
d: [2,174,1,2,1,3,21,173,130]
a: [3,178,1,2,1,174]
s: 178
i: [1,76,0,8,4,0,0,255] cc1_lattices
s: 76
d: [2,193,1,6,3,1,40,76,115]
d: [2,195,1,4,2,2,21,193,125]
d: [2,196,1,2,1,3,21,195,130]
a: [3,200,1,2,1,196]
s: 200
i: [1,77,0,8,4,0,0,255] cc1_lattices
s: 77
i: [1,78,0,8,4,0,0,255] cc1_lattices
s: 78
d: [2,207,1,6,3,1,40,77,115]
d: [2,209,1,4,2,2,21,207,125]
d: [2,210,1,2,1,3,21,209,130]
a: [3,214,1,2,1,210]
s: 214
i: [1,79,0,8,4,0,0,255] cc1_lattices
s: 79
d: [2,221,1,6,3,1,40,78,115]
d: [2,223,1,4,2,2,21,221,125]
d: [2,224,1,2,1,3,21,223,130]
a: [3,228,1,2,1,224]
s: 228
i: [1,80,0,8,4,0,0,255] cc1_lattices
s: 80
i: [1,82,0,9,3,0,0,1] dt_k5_lattices
s: 82
d: [2,235,1,6,3,1,40,79,115]
d: [2,237,1,4,2,2,21,235,125]
d: [2,238,1,2,1,3,21,237,130]
a: [3,242,1,2,1,238]
s: 242
i: [1,95,0,18,6,0,0,255] redefinition_r1_boolealg
d: [2,96,0,14,5,0,18,95]
s: 96
d: [2,249,1,6,3,1,40,80,115]
d: [2,251,1,4,2,2,21,249,125]
d: [2,252,1,2,1,3,21,251,130]
a: [3,256,1,2,1,252]
s: 256
i: [1,83,0,17,6,0,0,1] t16_lattices
s: 83
i: [1,73,0,18,9,0,0,255] cc2_lattices
s: 73
d: [2,263,1,10,4,1,40,96,110]
d: [2,267,1,8,3,2,21,263,115]
d: [2,268,1,6,2,3,21,267,125]
d: [2,269,1,4,1,4,21,268,130]
a: [3,273,1,4,1,269]
s: 273
i: [1,88,0,18,6,0,0,1] reflexivity_r1_boolealg
s: 88
d: [2,148,1,2,1,1,40,70,145]
a: [3,157,1,2,1,148]
s: 157
i: [1,90,0,21,7,0,0,255] redefinition_r1_boolealg
s: 90
i: [1,85,0,22,8,0,0,1] reflexivity_r3_lattices
s: 85
d: [2,279,1,18,7,1,40,85,110]
d: [2,283,1,16,6,2,21,279,115]
d: [2,284,1,14,5,3,21,283,256]
d: [2,285,1,12,4,4,21,284,242]
d: [2,286,1,10,3,5,21,285,214]
d: [2,287,1,8,2,6,21,286,130]
s: 287
i: [1,89,0,22,7,0,0,1] symmetry_r1_boolealg
s: 89
d: [2,289,2,4,1,6,40,287,110]
a: [3,295,2,4,1,289]
s: 295
i: [1,81,0,25,8,0,0,255] t8_lattices
s: 81
i: [1,86,0,26,9,0,0,255] redefinition_r3_lattices
s: 86
d: [2,312,3,22,8,1,40,86,295]
d: [2,315,3,18,7,1,18,312]
d: [2,331,3,14,6,2,21,315,110]
d: [2,332,3,12,5,3,21,331,115]
d: [2,333,3,10,4,4,21,332,256]
d: [2,334,3,8,3,5,21,333,242]
d: [2,335,3,6,2,6,21,334,214]
d: [2,336,3,4,1,7,21,335,130]
a: [3,340,3,4,1,336]
s: 340
i: [1,87,0,26,9,0,0,255] redefinition_r3_lattices
s: 87
d: [2,311,1,24,8,1,40,86,105]
d: [2,316,1,20,7,2,21,311,110]
d: [2,317,1,18,6,3,21,316,115]
d: [2,318,1,16,5,4,21,317,256]
d: [2,319,1,14,4,5,21,318,242]
d: [2,320,1,12,3,6,21,319,214]
d: [2,321,1,10,2,7,21,320,130]
a: [3,325,1,5,1,321]
s: 325
d: [2,349,2,0,0,3,52,130,152,325,82]
e: 349
d: [2,350,2,4,2,1,40,325,82]
d: [2,352,2,2,1,2,21,350,152]
d: [2,353,2,0,0,3,21,352,130]
e: 353
r: 325
a: [3,324,1,5,1,321]
s: 324
a: [3,329,1,5,1,321]
s: 329
d: [2,368,2,24,7,1,40,329,81]
d: [2,369,2,20,6,2,21,368,110]
d: [2,370,2,15,5,3,21,369,324]
d: [2,371,2,13,4,4,21,370,164]
d: [2,372,2,11,3,5,21,371,178]
d: [2,373,2,9,2,6,21,372,130]
a: [3,377,2,5,1,373]
s: 377
d: [2,160,1,2,1,1,40,71,145]
a: [3,169,1,2,1,160]
s: 169
d: [2,291,2,10,3,6,40,287,82]
d: [2,302,2,8,2,7,21,291,152]
d: [2,303,2,6,1,7,21,302,130]
a: [3,307,2,6,1,303]
s: 307
d: [2,172,1,6,3,1,40,75,145]
a: [3,191,1,2,1,172]
s: 191
d: [2,265,1,16,6,0,40,96,82]
d: [2,266,1,14,5,0,18,265]
d: [2,275,1,12,4,0,21,266,70]
s: 275
d: [2,357,2,12,4,1,40,324,96]
d: [2,359,2,10,3,2,21,357,115]
d: [2,360,2,8,2,3,21,359,125]
d: [2,361,2,6,1,4,21,360,130]
a: [3,365,2,6,1,361]
s: 365
d: [2,264,1,14,4,0,40,96,84]
s: 264
d: [2,383,3,26,8,1,40,307,86]
d: [2,384,3,21,7,1,18,383]
d: [2,385,3,16,6,2,21,384,324]
d: [2,386,3,14,5,3,21,385,115]
d: [2,387,3,12,4,4,21,386,256]
d: [2,388,3,10,3,5,21,387,242]
d: [2,389,3,8,2,6,21,388,214]
d: [2,390,3,6,1,7,21,389,130]
a: [3,394,3,6,1,390]
s: 394
d: [2,280,1,18,7,0,40,85,84]
s: 280
r: 85
d: [2,281,1,22,9,0,40,85,82]
d: [2,282,1,20,8,0,18,281]
d: [2,288,1,18,7,0,21,282,70]
s: 288
d: [2,290,2,8,1,6,40,287,84]
a: [3,300,2,8,1,290]
s: 300
d: [2,313,1,36,13,0,40,86,83]
d: [2,314,1,28,10,0,18,313]
d: [2,342,1,26,9,0,21,314,77]
d: [2,343,1,24,8,0,21,342,79]
d: [2,344,1,22,7,0,21,343,80]
s: 344
d: [2,431,3,0,0,7,52,125,120,130,115,110,377,324,344]
e: 431
f: 431,373,353,321,269,252,238,210,174,159,147,63,64,65,66,67,68,69
% Refutation found. Thanks to Tanya!
% SZS status Theorem for boolealg__t9_boolealg
% SZS output start Proof for boolealg__t9_boolealg
1. ! [X0] : ((l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => (r3_lattices(X0,X1,k5_lattices(X0)) => r1_boolealg(X0,X1,k5_lattices(X0))))) [input t9_boolealg]
2. ~! [X0] : ((l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => (r3_lattices(X0,X1,k5_lattices(X0)) => r1_boolealg(X0,X1,k5_lattices(X0))))) [negated conjecture 1]
3. ! [X0] : (l3_lattices(X0) => ((v10_lattices(X0) & ~v2_struct_0(X0)) => (v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)))) [input cc1_lattices]
5. ! [X0] : ((l1_lattices(X0) & ~v2_struct_0(X0)) => m1_subset_1(k5_lattices(X0),u1_struct_0(X0))) [input dt_k5_lattices]
9. ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))) [input dt_l3_lattices]
17. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (r1_boolealg(X0,X1,X2) <=> X1 = X2)) [input redefinition_r1_boolealg]
18. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l3_lattices(X0) & v9_lattices(X0) & v8_lattices(X0) & v6_lattices(X0) & ~v2_struct_0(X0)) => (r3_lattices(X0,X1,X2) <=> r1_lattices(X0,X1,X2))) [input redefinition_r3_lattices]
22. ! [X0] : ((l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => r3_lattices(X0,k5_lattices(X0),X1))) [input t16_lattices]
23. ! [X0] : ((l2_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ((r1_lattices(X0,X2,X1) & r1_lattices(X0,X1,X2)) => X1 = X2)))) [input t8_lattices]
27. ? [X0] : (? [X1] : ((~r1_boolealg(X0,X1,k5_lattices(X0)) & r3_lattices(X0,X1,k5_lattices(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
28. ? [X0] : (? [X1] : (~r1_boolealg(X0,X1,k5_lattices(X0)) & r3_lattices(X0,X1,k5_lattices(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) [flattening 27]
29. ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~l3_lattices(X0)) [ennf transformation 9]
32. ! [X0] : (((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)) | (~v10_lattices(X0) | v2_struct_0(X0))) | ~l3_lattices(X0)) [ennf transformation 3]
33. ! [X0] : ((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~v2_struct_0(X0)) | ~v10_lattices(X0) | v2_struct_0(X0) | ~l3_lattices(X0)) [flattening 32]
34. ! [X0] : (! [X1] : (! [X2] : ((X1 = X2 | (~r1_lattices(X0,X2,X1) | ~r1_lattices(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_lattices(X0) | ~v4_lattices(X0) | v2_struct_0(X0))) [ennf transformation 23]
35. ! [X0] : (! [X1] : (! [X2] : (X1 = X2 | ~r1_lattices(X0,X2,X1) | ~r1_lattices(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_lattices(X0) | ~v4_lattices(X0) | v2_struct_0(X0)) [flattening 34]
36. ! [X0] : (m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | (~l1_lattices(X0) | v2_struct_0(X0))) [ennf transformation 5]
37. ! [X0] : (m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | ~l1_lattices(X0) | v2_struct_0(X0)) [flattening 36]
38. ! [X0] : (! [X1] : (r3_lattices(X0,k5_lattices(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 22]
39. ! [X0] : (! [X1] : (r3_lattices(X0,k5_lattices(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 38]
42. ! [X0,X1,X2] : ((r3_lattices(X0,X1,X2) <=> r1_lattices(X0,X1,X2)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v9_lattices(X0) | ~v8_lattices(X0) | ~v6_lattices(X0) | v2_struct_0(X0))) [ennf transformation 18]
43. ! [X0,X1,X2] : ((r3_lattices(X0,X1,X2) <=> r1_lattices(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v9_lattices(X0) | ~v8_lattices(X0) | ~v6_lattices(X0) | v2_struct_0(X0)) [flattening 42]
48. ! [X0,X1,X2] : ((r1_boolealg(X0,X1,X2) <=> X1 = X2) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 17]
49. ! [X0,X1,X2] : ((r1_boolealg(X0,X1,X2) <=> X1 = X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 48]
50. ? [X0] : (? [X1] : (~r1_boolealg(X0,X1,k5_lattices(X0)) & r3_lattices(X0,X1,k5_lattices(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (? [X1] : (~r1_boolealg(sK0,X1,k5_lattices(sK0)) & r3_lattices(sK0,X1,k5_lattices(sK0)) & m1_subset_1(X1,u1_struct_0(sK0))) & l3_lattices(sK0) & v13_lattices(sK0) & v10_lattices(sK0) & ~v2_struct_0(sK0)) [choice axiom]
51. ? [X1] : (~r1_boolealg(sK0,X1,k5_lattices(sK0)) & r3_lattices(sK0,X1,k5_lattices(sK0)) & m1_subset_1(X1,u1_struct_0(sK0))) => (~r1_boolealg(sK0,sK1,k5_lattices(sK0)) & r3_lattices(sK0,sK1,k5_lattices(sK0)) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
52. (~r1_boolealg(sK0,sK1,k5_lattices(sK0)) & r3_lattices(sK0,sK1,k5_lattices(sK0)) & m1_subset_1(sK1,u1_struct_0(sK0))) & l3_lattices(sK0) & v13_lattices(sK0) & v10_lattices(sK0) & ~v2_struct_0(sK0) [skolemisation 28,51,50]
55. ! [X0,X1,X2] : (((r3_lattices(X0,X1,X2) | ~r1_lattices(X0,X1,X2)) & (r1_lattices(X0,X1,X2) | ~r3_lattices(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v9_lattices(X0) | ~v8_lattices(X0) | ~v6_lattices(X0) | v2_struct_0(X0)) [nnf transformation 43]
56. ! [X0,X1,X2] : (((r1_boolealg(X0,X1,X2) | X1 != X2) & (X1 = X2 | ~r1_boolealg(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [nnf transformation 49]
63. ~v2_struct_0(sK0) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
64. v10_lattices(sK0) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. v13_lattices(sK0) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
66. l3_lattices(sK0) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
67. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 52] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
68. r3_lattices(sK0,sK1,k5_lattices(sK0)) [cnf transformation 52] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
69. ~r1_boolealg(sK0,sK1,k5_lattices(sK0)) [cnf transformation 52] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
70. ~l3_lattices(X0) | l1_lattices(X0) [cnf transformation 29] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
71. ~l3_lattices(X0) | l2_lattices(X0) [cnf transformation 29] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
75. ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | v4_lattices(X0) [cnf transformation 33] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
77. ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | v6_lattices(X0) [cnf transformation 33] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
79. ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | v8_lattices(X0) [cnf transformation 33] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
80. ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) | v9_lattices(X0) [cnf transformation 33] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
81. ~r1_lattices(X0,X2,X1) | X1 = X2 | ~r1_lattices(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_lattices(X0) | ~v4_lattices(X0) | v2_struct_0(X0) [cnf transformation 35] {a:0,w:25,wCS:10000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
82. m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | ~l1_lattices(X0) | v2_struct_0(X0) [cnf transformation 37] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
83. r3_lattices(X0,k5_lattices(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 39] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
86. ~r3_lattices(X0,X1,X2) | r1_lattices(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v9_lattices(X0) | ~v8_lattices(X0) | ~v6_lattices(X0) | v2_struct_0(X0) [cnf transformation 55] {a:0,w:26,wCS:10400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
91. r1_boolealg(X0,X1,X2) | X1 != X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 56] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:1,thDist:-1}
95. r1_boolealg(X0,X2,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [equality resolution 91] {a:0,w:18,wCS:7200,sine:255,thAx:0,allAx:1,thDist:-1}
96. ~m1_subset_1(X2,u1_struct_0(X0)) | r1_boolealg(X0,X2,X2) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [duplicate literal removal 95] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
98. 1 <=> r1_boolealg(sK0,sK1,k5_lattices(sK0)) [avatar definition]
101. ~1 [avatar split clause 69,98]
103. 2 <=> r3_lattices(sK0,sK1,k5_lattices(sK0)) [avatar definition]
105. r3_lattices(sK0,sK1,k5_lattices(sK0)) <- (2) [avatar component clause 103] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
106. 2 [avatar split clause 68,103]
108. 3 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
110. m1_subset_1(sK1,u1_struct_0(sK0)) <- (3) [avatar component clause 108] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
111. 3 [avatar split clause 67,108]
113. 4 <=> l3_lattices(sK0) [avatar definition]
115. l3_lattices(sK0) <- (4) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. 4 [avatar split clause 66,113]
118. 5 <=> v13_lattices(sK0) [avatar definition]
120. v13_lattices(sK0) <- (5) [avatar component clause 118] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
121. 5 [avatar split clause 65,118]
123. 6 <=> v10_lattices(sK0) [avatar definition]
125. v10_lattices(sK0) <- (6) [avatar component clause 123] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. 6 [avatar split clause 64,123]
128. 7 <=> v2_struct_0(sK0) [avatar definition]
130. ~v2_struct_0(sK0) <- (~7) [avatar component clause 128] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
131. ~7 [avatar split clause 63,128]
147. l1_lattices(sK0) <- (4) [resolution 70,115] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
150. 11 <=> l1_lattices(sK0) [avatar definition]
152. l1_lattices(sK0) <- (11) [avatar component clause 150] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
153. 11 | ~4 [avatar split clause 147,113,150]
159. l2_lattices(sK0) <- (4) [resolution 71,115] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
162. 13 <=> l2_lattices(sK0) [avatar definition]
164. l2_lattices(sK0) <- (13) [avatar component clause 162] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
165. 13 | ~4 [avatar split clause 159,113,162]
171. ~v10_lattices(sK0) | v2_struct_0(sK0) | v4_lattices(sK0) <- (4) [resolution 75,115] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
173. v2_struct_0(sK0) | v4_lattices(sK0) <- (4, 6) [subsumption resolution 171,125] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
174. v4_lattices(sK0) <- (4, 6, ~7) [subsumption resolution 173,130] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
176. 15 <=> v4_lattices(sK0) [avatar definition]
178. v4_lattices(sK0) <- (15) [avatar component clause 176] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
179. 15 | ~4 | ~6 | 7 [avatar split clause 174,128,123,113,176]
207. ~v10_lattices(sK0) | v2_struct_0(sK0) | v6_lattices(sK0) <- (4) [resolution 77,115] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
209. v2_struct_0(sK0) | v6_lattices(sK0) <- (4, 6) [subsumption resolution 207,125] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
210. v6_lattices(sK0) <- (4, 6, ~7) [subsumption resolution 209,130] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
212. 21 <=> v6_lattices(sK0) [avatar definition]
214. v6_lattices(sK0) <- (21) [avatar component clause 212] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
215. 21 | ~4 | ~6 | 7 [avatar split clause 210,128,123,113,212]
235. ~v10_lattices(sK0) | v2_struct_0(sK0) | v8_lattices(sK0) <- (4) [resolution 79,115] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
237. v2_struct_0(sK0) | v8_lattices(sK0) <- (4, 6) [subsumption resolution 235,125] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
238. v8_lattices(sK0) <- (4, 6, ~7) [subsumption resolution 237,130] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
240. 25 <=> v8_lattices(sK0) [avatar definition]
242. v8_lattices(sK0) <- (25) [avatar component clause 240] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
243. 25 | ~4 | ~6 | 7 [avatar split clause 238,128,123,113,240]
249. ~v10_lattices(sK0) | v2_struct_0(sK0) | v9_lattices(sK0) <- (4) [resolution 80,115] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
251. v2_struct_0(sK0) | v9_lattices(sK0) <- (4, 6) [subsumption resolution 249,125] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
252. v9_lattices(sK0) <- (4, 6, ~7) [subsumption resolution 251,130] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
254. 27 <=> v9_lattices(sK0) [avatar definition]
256. v9_lattices(sK0) <- (27) [avatar component clause 254] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
257. 27 | ~4 | ~6 | 7 [avatar split clause 252,128,123,113,254]
263. r1_boolealg(sK0,sK1,sK1) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (3) [resolution 96,110] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
267. r1_boolealg(sK0,sK1,sK1) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (3, 4) [subsumption resolution 263,115] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
268. r1_boolealg(sK0,sK1,sK1) | v2_struct_0(sK0) <- (3, 4, 6) [subsumption resolution 267,125] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
269. r1_boolealg(sK0,sK1,sK1) <- (3, 4, 6, ~7) [subsumption resolution 268,130] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
271. 29 <=> r1_boolealg(sK0,sK1,sK1) [avatar definition]
274. 29 | ~3 | ~4 | ~6 | 7 [avatar split clause 269,128,123,113,108,271]
311. r1_lattices(sK0,sK1,k5_lattices(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v9_lattices(sK0) | ~v8_lattices(sK0) | ~v6_lattices(sK0) | v2_struct_0(sK0) <- (2) [resolution 86,105] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
313. r1_lattices(X0,k5_lattices(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | ~l3_lattices(X0) | ~v9_lattices(X0) | ~v8_lattices(X0) | ~v6_lattices(X0) | v2_struct_0(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [resolution 86,83] {a:1,w:36,wCS:14400,sine:1,thAx:0,allAx:2,thDist:-2}
314. r1_lattices(X0,k5_lattices(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | ~l3_lattices(X0) | ~v9_lattices(X0) | ~v8_lattices(X0) | ~v6_lattices(X0) | v2_struct_0(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) [duplicate literal removal 313] {a:1,w:28,wCS:11200,sine:1,thAx:0,allAx:2,thDist:-2}
316. r1_lattices(sK0,sK1,k5_lattices(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v9_lattices(sK0) | ~v8_lattices(sK0) | ~v6_lattices(sK0) | v2_struct_0(sK0) <- (2, 3) [subsumption resolution 311,110] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
317. r1_lattices(sK0,sK1,k5_lattices(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~v9_lattices(sK0) | ~v8_lattices(sK0) | ~v6_lattices(sK0) | v2_struct_0(sK0) <- (2, 3, 4) [subsumption resolution 316,115] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
318. r1_lattices(sK0,sK1,k5_lattices(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~v8_lattices(sK0) | ~v6_lattices(sK0) | v2_struct_0(sK0) <- (2, 3, 4, 27) [subsumption resolution 317,256] {a:1,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
319. r1_lattices(sK0,sK1,k5_lattices(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~v6_lattices(sK0) | v2_struct_0(sK0) <- (2, 3, 4, 25, 27) [subsumption resolution 318,242] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
320. r1_lattices(sK0,sK1,k5_lattices(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | v2_struct_0(sK0) <- (2, 3, 4, 21, 25, 27) [subsumption resolution 319,214] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
321. r1_lattices(sK0,sK1,k5_lattices(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) <- (2, 3, 4, ~7, 21, 25, 27) [subsumption resolution 320,130] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
323. 33 <=> m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) [avatar definition]
324. m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) <- (33) [avatar component clause 323] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
325. ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) <- (~33) [avatar component clause 323] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
327. 34 <=> r1_lattices(sK0,sK1,k5_lattices(sK0)) [avatar definition]
329. r1_lattices(sK0,sK1,k5_lattices(sK0)) <- (34) [avatar component clause 327] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
330. ~33 | 34 | ~2 | ~3 | ~4 | 7 | ~21 | ~25 | ~27 [avatar split clause 321,254,240,212,128,113,108,103,327,323]
342. r1_lattices(X0,k5_lattices(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | ~l3_lattices(X0) | ~v9_lattices(X0) | ~v8_lattices(X0) | v2_struct_0(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) [subsumption resolution 314,77] {a:1,w:26,wCS:10400,sine:1,thAx:0,allAx:2,thDist:-2}
343. r1_lattices(X0,k5_lattices(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | ~l3_lattices(X0) | ~v9_lattices(X0) | v2_struct_0(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) [subsumption resolution 342,79] {a:1,w:24,wCS:9600,sine:1,thAx:0,allAx:2,thDist:-2}
344. ~m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | r1_lattices(X0,k5_lattices(X0),X1) | ~l3_lattices(X0) | v2_struct_0(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) [subsumption resolution 343,80] {a:1,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
350. ~l1_lattices(sK0) | v2_struct_0(sK0) <- (~33) [resolution 325,82] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
352. v2_struct_0(sK0) <- (11, ~33) [subsumption resolution 350,152] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
353. $false <- (~7, 11, ~33) [subsumption resolution 352,130] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
354. 7 | ~11 | 33 [avatar contradiction clause 353]
368. k5_lattices(sK0) = sK1 | ~r1_lattices(sK0,k5_lattices(sK0),sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~l2_lattices(sK0) | ~v4_lattices(sK0) | v2_struct_0(sK0) <- (34) [resolution 329,81] {a:2,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
369. k5_lattices(sK0) = sK1 | ~r1_lattices(sK0,k5_lattices(sK0),sK1) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~l2_lattices(sK0) | ~v4_lattices(sK0) | v2_struct_0(sK0) <- (3, 34) [subsumption resolution 368,110] {a:2,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
370. k5_lattices(sK0) = sK1 | ~r1_lattices(sK0,k5_lattices(sK0),sK1) | ~l2_lattices(sK0) | ~v4_lattices(sK0) | v2_struct_0(sK0) <- (3, 33, 34) [subsumption resolution 369,324] {a:2,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
371. k5_lattices(sK0) = sK1 | ~r1_lattices(sK0,k5_lattices(sK0),sK1) | ~v4_lattices(sK0) | v2_struct_0(sK0) <- (3, 13, 33, 34) [subsumption resolution 370,164] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
372. k5_lattices(sK0) = sK1 | ~r1_lattices(sK0,k5_lattices(sK0),sK1) | v2_struct_0(sK0) <- (3, 13, 15, 33, 34) [subsumption resolution 371,178] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
373. k5_lattices(sK0) = sK1 | ~r1_lattices(sK0,k5_lattices(sK0),sK1) <- (3, ~7, 13, 15, 33, 34) [subsumption resolution 372,130] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
375. 37 <=> r1_lattices(sK0,k5_lattices(sK0),sK1) [avatar definition]
377. ~r1_lattices(sK0,k5_lattices(sK0),sK1) <- (~37) [avatar component clause 375] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
379. 38 <=> k5_lattices(sK0) = sK1 [avatar definition]
382. ~37 | 38 | ~3 | 7 | ~13 | ~15 | ~33 | ~34 [avatar split clause 373,327,323,176,162,128,108,379,375]
431. $false <- (3, 4, 5, 6, ~7, 33, ~37) [unit resulting resolution 125,120,130,115,110,377,324,344] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:26,thDist:-26}
435. ~3 | ~4 | ~5 | ~6 | 7 | ~33 | 37 [avatar contradiction clause 431]
441. k5_lattices(sK0) != sK1 | r1_boolealg(sK0,sK1,k5_lattices(sK0)) | ~r1_boolealg(sK0,sK1,sK1) [theory tautology sat conflict] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:1,thDist:-1}
442. $false [avatar sat refutation 101,106,111,116,121,126,131,153,165,179,215,243,257,274,330,354,382,435,441] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:338,thDist:-338}
% SZS output end Proof for boolealg__t9_boolealg
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 32
% Generated clauses: 210
% Active clauses: 56
% Passive clauses: 61
% Final active clauses: 54
% Final passive clauses: 3
% Main loop iterations started: 58
% 
% >>> Simplifying Inferences
% Duplicate literals: 35
% Fw subsumption resolutions: 67
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 6
% Forward subsumptions: 13
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 55
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 22
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 43
% SAT solver unit clauses: 10
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9594
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% splitting component index maintenance: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

