i: [1,49,0,7,2,1,0,0] t35_int_1
a: [3,70,0,2,1,49]
s: 70
i: [1,45,0,2,1,1,0,0] t35_int_1
a: [3,81,0,2,1,45]
s: 81
a: [3,75,0,5,1,49]
s: 75
i: [1,63,0,7,3,0,0,255] reflexivity_r1_xxreal_0
a: [3,85,0,2,1,63]
s: 85
i: [1,65,0,2,1,0,0,255] rc1_xxreal_0
a: [3,93,0,2,1,65]
s: 93
d: [2,105,0,0,0,2,21,93,85]
e: 105
r: 85
s: 93
i: [1,66,0,2,1,0,0,1] rc1_xreal_0
a: [3,98,0,2,1,66]
s: 98
i: [1,67,0,2,1,0,0,1] rc2_int_1
a: [3,103,0,2,1,67]
s: 103
i: [1,51,0,4,2,0,0,255] cc4_xreal_0
s: 51
i: [1,62,0,4,2,0,0,1] cc3_int_1
s: 62
i: [1,52,0,5,2,0,0,1] dt_k2_int_1
s: 52
d: [2,107,1,2,1,1,40,51,81]
a: [3,112,1,2,1,107]
s: 112
i: [1,53,0,5,2,0,0,1] dt_k1_int_1
s: 53
d: [2,108,1,2,1,1,40,51,98]
a: [3,117,1,2,1,108]
s: 117
a: [3,88,0,5,2,63]
s: 88
i: [1,54,0,8,2,0,0,255] projectivity_k1_int_1
s: 54
d: [2,119,1,2,1,1,40,62,103]
a: [3,123,1,2,1,119]
s: 123
i: [1,55,0,8,2,0,0,255] projectivity_k2_int_1
s: 55
d: [2,129,2,2,1,1,40,123,51]
a: [3,133,2,2,1,129]
s: 133
i: [1,56,0,8,3,0,0,1] t26_int_1
s: 56
d: [2,137,0,6,2,0,21,56,62]
s: 137
d: [2,125,1,5,2,0,40,52,62]
s: 125
i: [1,57,0,8,3,0,0,1] t26_int_1
s: 57
d: [2,126,1,5,2,0,40,53,62]
s: 126
i: [1,58,0,8,3,0,0,1] t31_int_1
s: 58
d: [2,146,0,6,2,0,21,58,62]
s: 146
d: [2,140,2,5,2,0,40,125,51]
s: 140
i: [1,59,0,8,3,0,0,1] t31_int_1
s: 59
d: [2,144,2,5,2,0,40,126,51]
s: 144
i: [1,60,0,9,3,0,0,255] t34_int_1
s: 60
i: [1,61,0,9,3,0,0,255] t34_int_1
s: 61
d: [2,154,0,7,2,0,21,61,62]
s: 154
i: [1,64,0,10,4,0,0,255] connectedness_r1_xxreal_0
s: 64
d: [2,162,1,8,3,0,44,137,154]
d: [2,165,1,6,2,0,18,162]
s: 165
i: [1,50,0,15,6,0,0,255] t2_xxreal_0
s: 50
d: [2,179,1,16,5,1,40,50,75]
a: [3,190,1,3,1,179]
s: 190
d: [2,199,3,0,0,2,52,81,190,140]
e: 199
d: [2,200,3,2,1,1,40,190,140]
d: [2,202,3,0,0,2,21,200,81]
e: 202
r: 190
a: [3,189,1,3,1,179]
s: 189
a: [3,194,1,3,1,179]
s: 194
d: [2,204,3,0,0,2,52,81,194,144]
e: 204
d: [2,205,3,2,1,1,40,194,144]
d: [2,208,3,0,0,2,21,205,81]
e: 208
r: 194
a: [3,193,1,3,1,179]
s: 193
a: [3,197,1,10,3,179]
s: 197
d: [2,138,1,10,3,0,44,137,54]
d: [2,139,1,7,2,0,21,138,53]
s: 139
d: [2,211,2,13,3,1,40,197,146]
a: [3,220,2,3,1,211]
s: 220
d: [2,236,3,0,0,2,52,81,220,52]
e: 236
d: [2,237,3,2,1,1,40,220,52]
d: [2,239,3,0,0,2,21,237,81]
e: 239
r: 220
d: [2,147,1,10,3,0,44,146,55]
d: [2,148,1,7,2,0,21,147,52]
s: 148
a: [3,219,2,3,1,211]
s: 219
d: [2,155,1,11,3,0,44,154,54]
d: [2,166,1,8,2,0,21,155,53]
s: 166
d: [2,158,1,10,3,0,44,54,154]
d: [2,167,1,8,2,0,21,158,62]
s: 167
d: [2,245,3,3,1,1,40,219,62]
a: [3,249,3,3,1,245]
s: 249
d: [2,164,1,11,3,0,44,54,154]
d: [2,168,1,8,2,0,21,164,53]
s: 168
a: [3,228,2,4,1,211]
s: 228
d: [2,276,4,0,0,2,52,249,228,140]
e: 276
d: [2,277,3,3,1,1,40,228,140]
d: [2,280,3,0,0,2,21,277,249]
e: 280
d: [2,278,3,5,2,1,44,228,55]
d: [2,282,3,2,1,2,21,278,189]
d: [2,283,3,0,0,3,21,282,81]
e: 283
r: 70
r: 75
r: 197
r: 228
a: [3,71,0,2,1,49]
s: 71
a: [3,74,0,5,1,49]
s: 74
a: [3,227,2,4,1,211]
s: 227
d: [2,286,1,11,3,1,40,74,64]
d: [2,288,1,8,2,2,21,286,189]
d: [2,289,1,5,1,3,21,288,193]
a: [3,293,1,5,1,289]
s: 293
r: 71
r: 74
s: 70
s: 75
d: [2,287,1,7,2,1,44,74,154]
d: [2,295,1,5,1,2,21,287,71]
a: [3,299,1,5,1,295]
s: 299
d: [2,314,2,0,0,2,52,81,299,148]
e: 314
d: [2,315,2,0,0,3,52,189,299,88]
e: 315
d: [2,316,2,0,0,5,52,189,193,189,75,293,299,50]
e: 316
d: [2,317,2,0,0,2,52,189,189,299,299,64]
e: 317
d: [2,318,2,0,0,2,52,189,189,299,299,64]
e: 318
d: [2,319,2,2,1,1,40,299,148]
d: [2,328,2,0,0,2,21,319,81]
e: 328
d: [2,320,2,3,1,2,40,299,88]
d: [2,330,2,0,0,3,21,320,189]
e: 330
r: 299
a: [3,298,1,5,1,295]
s: 298
s: 197
a: [3,224,2,6,1,211]
s: 224
d: [2,169,1,13,5,0,40,64,57]
d: [2,173,1,11,4,0,21,169,51]
d: [2,174,1,8,3,0,21,173,144]
s: 174
d: [2,170,1,13,5,0,40,64,59]
d: [2,175,1,11,4,0,21,170,51]
d: [2,176,1,8,3,0,21,175,140]
s: 176
d: [2,302,2,16,5,1,40,293,50]
d: [2,304,2,13,4,2,21,302,189]
d: [2,305,2,10,3,3,21,304,193]
s: 305
d: [2,359,1,6,2,0,21,176,165]
s: 359
d: [2,367,3,16,4,3,40,305,174]
d: [2,391,3,12,3,3,21,367,144]
a: [3,395,3,3,1,391]
s: 395
d: [2,414,4,0,0,2,52,81,395,126]
e: 414
d: [2,415,4,2,1,1,40,395,126]
d: [2,418,4,0,0,2,21,415,81]
e: 418
r: 395
d: [2,182,1,19,8,0,40,50,64]
d: [2,185,1,15,6,0,18,182]
s: 185
d: [2,180,1,16,6,0,40,50,146]
s: 180
d: [2,362,3,13,3,3,40,305,146]
a: [3,379,3,3,1,362]
s: 379
d: [2,184,1,16,6,0,40,50,137]
s: 184
a: [3,394,3,3,1,391]
s: 394
d: [2,401,2,8,3,1,40,359,197]
d: [2,405,2,6,2,2,21,401,81]
d: [2,406,2,4,1,3,21,405,112]
a: [3,410,2,4,1,406]
s: 410
d: [2,491,3,0,0,3,52,81,70,410,57]
e: 491
d: [2,493,3,4,2,1,40,410,57]
d: [2,499,3,2,1,2,21,493,70]
d: [2,500,3,0,0,3,21,499,81]
e: 500
i: [1,46,0,7,2,1,0,0] t35_int_1
d: [2,309,1,16,5,1,40,75,50]
d: [2,311,1,13,4,2,21,309,193]
d: [2,312,1,10,3,3,21,311,189]
f: 500,406,328,312,295,208,202,107,45,46,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for int_1__t35_int_1
% SZS output start Proof for int_1__t35_int_1
1. ! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(k2_int_1(X0),k1_int_1(X0)) & ~v1_int_1(X0)) & ~(v1_int_1(X0) & ~r1_xxreal_0(k2_int_1(X0),k1_int_1(X0))))) [input t35_int_1]
2. ~! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(k2_int_1(X0),k1_int_1(X0)) & ~v1_int_1(X0)) & ~(v1_int_1(X0) & ~r1_xxreal_0(k2_int_1(X0),k1_int_1(X0))))) [negated conjecture 1]
3. ! [X0] : (v1_int_1(X0) => v1_xreal_0(X0)) [input cc3_int_1]
5. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
6. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
7. ! [X0] : (v1_xreal_0(X0) => v1_int_1(k1_int_1(X0))) [input dt_k1_int_1]
8. ! [X0] : (v1_xreal_0(X0) => v1_int_1(k2_int_1(X0))) [input dt_k2_int_1]
10. ! [X0] : (v1_xreal_0(X0) => k2_int_1(X0) = k2_int_1(k2_int_1(X0))) [input projectivity_k2_int_1]
15. ! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(X0,k1_int_1(X0)) & ~v1_int_1(X0)) & ~(v1_int_1(X0) & ~r1_xxreal_0(X0,k1_int_1(X0))))) [input t26_int_1]
16. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input t2_xxreal_0]
17. ! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(k2_int_1(X0),X0) & ~v1_int_1(X0)) & ~(v1_int_1(X0) & ~r1_xxreal_0(k2_int_1(X0),X0)))) [input t31_int_1]
18. ! [X0] : (v1_xreal_0(X0) => (k2_int_1(X0) = k1_int_1(X0) <=> v1_int_1(X0))) [input t34_int_1]
20. ? [X0] : (((r1_xxreal_0(k2_int_1(X0),k1_int_1(X0)) & ~v1_int_1(X0)) | (v1_int_1(X0) & ~r1_xxreal_0(k2_int_1(X0),k1_int_1(X0)))) & v1_xreal_0(X0)) [ennf transformation 2]
21. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 16]
22. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 21]
23. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 5]
24. ! [X0] : (v1_int_1(k2_int_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 8]
25. ! [X0] : (v1_int_1(k1_int_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 7]
27. ! [X0] : (k2_int_1(X0) = k2_int_1(k2_int_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 10]
28. ! [X0] : (((~r1_xxreal_0(X0,k1_int_1(X0)) | v1_int_1(X0)) & (~v1_int_1(X0) | r1_xxreal_0(X0,k1_int_1(X0)))) | ~v1_xreal_0(X0)) [ennf transformation 15]
29. ! [X0] : (((~r1_xxreal_0(k2_int_1(X0),X0) | v1_int_1(X0)) & (~v1_int_1(X0) | r1_xxreal_0(k2_int_1(X0),X0))) | ~v1_xreal_0(X0)) [ennf transformation 17]
30. ! [X0] : ((k2_int_1(X0) = k1_int_1(X0) <=> v1_int_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 18]
31. ! [X0] : (v1_xreal_0(X0) | ~v1_int_1(X0)) [ennf transformation 3]
34. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 6]
35. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 34]
36. ? [X0] : (((r1_xxreal_0(k2_int_1(X0),k1_int_1(X0)) & ~v1_int_1(X0)) | (v1_int_1(X0) & ~r1_xxreal_0(k2_int_1(X0),k1_int_1(X0)))) & v1_xreal_0(X0)) => (((r1_xxreal_0(k2_int_1(sK0),k1_int_1(sK0)) & ~v1_int_1(sK0)) | (v1_int_1(sK0) & ~r1_xxreal_0(k2_int_1(sK0),k1_int_1(sK0)))) & v1_xreal_0(sK0)) [choice axiom]
37. ((r1_xxreal_0(k2_int_1(sK0),k1_int_1(sK0)) & ~v1_int_1(sK0)) | (v1_int_1(sK0) & ~r1_xxreal_0(k2_int_1(sK0),k1_int_1(sK0)))) & v1_xreal_0(sK0) [skolemisation 20,36]
38. ! [X0] : (((k2_int_1(X0) = k1_int_1(X0) | ~v1_int_1(X0)) & (v1_int_1(X0) | k2_int_1(X0) != k1_int_1(X0))) | ~v1_xreal_0(X0)) [nnf transformation 30]
45. v1_xreal_0(sK0) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. ~v1_int_1(sK0) | ~r1_xxreal_0(k2_int_1(sK0),k1_int_1(sK0)) [cnf transformation 37] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. r1_xxreal_0(k2_int_1(sK0),k1_int_1(sK0)) | v1_int_1(sK0) [cnf transformation 37] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 22] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
51. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 23] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
52. v1_int_1(k2_int_1(X0)) | ~v1_xreal_0(X0) [cnf transformation 24] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. v1_int_1(k1_int_1(X0)) | ~v1_xreal_0(X0) [cnf transformation 25] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. k2_int_1(X0) = k2_int_1(k2_int_1(X0)) | ~v1_xreal_0(X0) [cnf transformation 27] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
56. ~v1_int_1(X0) | r1_xxreal_0(X0,k1_int_1(X0)) | ~v1_xreal_0(X0) [cnf transformation 28] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
57. ~r1_xxreal_0(X0,k1_int_1(X0)) | v1_int_1(X0) | ~v1_xreal_0(X0) [cnf transformation 28] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. ~v1_int_1(X0) | r1_xxreal_0(k2_int_1(X0),X0) | ~v1_xreal_0(X0) [cnf transformation 29] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
59. ~r1_xxreal_0(k2_int_1(X0),X0) | v1_int_1(X0) | ~v1_xreal_0(X0) [cnf transformation 29] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. k2_int_1(X0) = k1_int_1(X0) | ~v1_int_1(X0) | ~v1_xreal_0(X0) [cnf transformation 38] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
62. ~v1_int_1(X0) | v1_xreal_0(X0) [cnf transformation 31] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 35] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
69. 1 <=> v1_int_1(sK0) [avatar definition]
70. ~v1_int_1(sK0) <- (~1) [avatar component clause 69] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. v1_int_1(sK0) <- (1) [avatar component clause 69] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
73. 2 <=> r1_xxreal_0(k2_int_1(sK0),k1_int_1(sK0)) [avatar definition]
74. ~r1_xxreal_0(k2_int_1(sK0),k1_int_1(sK0)) <- (~2) [avatar component clause 73] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. r1_xxreal_0(k2_int_1(sK0),k1_int_1(sK0)) <- (2) [avatar component clause 73] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. 1 | 2 [avatar split clause 49,73,69]
77. ~2 | ~1 [avatar split clause 46,69,73]
79. 3 <=> v1_xreal_0(sK0) [avatar definition]
81. v1_xreal_0(sK0) <- (3) [avatar component clause 79] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
82. 3 [avatar split clause 45,79]
107. v1_xxreal_0(sK0) <- (3) [resolution 51,81] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. 9 <=> v1_xxreal_0(sK0) [avatar definition]
112. v1_xxreal_0(sK0) <- (9) [avatar component clause 110] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. 9 | ~3 [avatar split clause 107,79,110]
125. v1_xreal_0(k2_int_1(X0)) | ~v1_xreal_0(X0) [resolution 52,62] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
126. v1_xreal_0(k1_int_1(X0)) | ~v1_xreal_0(X0) [resolution 53,62] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
137. r1_xxreal_0(X0,k1_int_1(X0)) | ~v1_int_1(X0) [subsumption resolution 56,62] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
140. v1_xxreal_0(k2_int_1(X0)) | ~v1_xreal_0(X0) [resolution 125,51] {a:2,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
144. v1_xxreal_0(k1_int_1(X0)) | ~v1_xreal_0(X0) [resolution 126,51] {a:2,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
146. r1_xxreal_0(k2_int_1(X0),X0) | ~v1_int_1(X0) [subsumption resolution 58,62] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
147. r1_xxreal_0(k2_int_1(X0),k2_int_1(X0)) | ~v1_int_1(k2_int_1(X0)) | ~v1_xreal_0(X0) [superposition 146,55] {a:1,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
148. r1_xxreal_0(k2_int_1(X0),k2_int_1(X0)) | ~v1_xreal_0(X0) [subsumption resolution 147,52] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
154. k2_int_1(X0) = k1_int_1(X0) | ~v1_int_1(X0) [subsumption resolution 61,62] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
162. r1_xxreal_0(X5,k2_int_1(X5)) | ~v1_int_1(X5) | ~v1_int_1(X5) [superposition 137,154] {a:1,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
165. r1_xxreal_0(X5,k2_int_1(X5)) | ~v1_int_1(X5) [duplicate literal removal 162] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
170. r1_xxreal_0(X1,k2_int_1(X1)) | ~v1_xxreal_0(k2_int_1(X1)) | ~v1_xxreal_0(X1) | v1_int_1(X1) | ~v1_xreal_0(X1) [resolution 64,59] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
175. r1_xxreal_0(X1,k2_int_1(X1)) | ~v1_xxreal_0(k2_int_1(X1)) | v1_int_1(X1) | ~v1_xreal_0(X1) [subsumption resolution 170,51] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
176. r1_xxreal_0(X1,k2_int_1(X1)) | v1_int_1(X1) | ~v1_xreal_0(X1) [subsumption resolution 175,140] {a:1,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
188. 13 <=> v1_xxreal_0(k2_int_1(sK0)) [avatar definition]
189. v1_xxreal_0(k2_int_1(sK0)) <- (13) [avatar component clause 188] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
190. ~v1_xxreal_0(k2_int_1(sK0)) <- (~13) [avatar component clause 188] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
192. 14 <=> v1_xxreal_0(k1_int_1(sK0)) [avatar definition]
193. v1_xxreal_0(k1_int_1(sK0)) <- (14) [avatar component clause 192] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
194. ~v1_xxreal_0(k1_int_1(sK0)) <- (~14) [avatar component clause 192] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
196. 15 <=> ! [X0] : (r1_xxreal_0(X0,k1_int_1(sK0)) | ~v1_xxreal_0(X0) | ~r1_xxreal_0(X0,k2_int_1(sK0))) [avatar definition]
197. ~r1_xxreal_0(X0,k2_int_1(sK0)) | ~v1_xxreal_0(X0) | r1_xxreal_0(X0,k1_int_1(sK0)) <- (15) [avatar component clause 196] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
200. ~v1_xreal_0(sK0) <- (~13) [resolution 190,140] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
202. $false <- (3, ~13) [subsumption resolution 200,81] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
203. ~3 | 13 [avatar contradiction clause 202]
205. ~v1_xreal_0(sK0) <- (~14) [resolution 194,144] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
208. $false <- (3, ~14) [subsumption resolution 205,81] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
209. ~3 | 14 [avatar contradiction clause 208]
287. ~r1_xxreal_0(k2_int_1(sK0),k2_int_1(sK0)) | ~v1_int_1(sK0) <- (~2) [superposition 74,154] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
295. ~r1_xxreal_0(k2_int_1(sK0),k2_int_1(sK0)) <- (1, ~2) [subsumption resolution 287,71] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
297. 21 <=> r1_xxreal_0(k2_int_1(sK0),k2_int_1(sK0)) [avatar definition]
299. ~r1_xxreal_0(k2_int_1(sK0),k2_int_1(sK0)) <- (~21) [avatar component clause 297] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
300. ~21 | ~1 | 2 [avatar split clause 295,73,69,297]
309. r1_xxreal_0(X0,k1_int_1(sK0)) | ~r1_xxreal_0(X0,k2_int_1(sK0)) | ~v1_xxreal_0(k1_int_1(sK0)) | ~v1_xxreal_0(k2_int_1(sK0)) | ~v1_xxreal_0(X0) <- (2) [resolution 75,50] {a:1,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
311. r1_xxreal_0(X0,k1_int_1(sK0)) | ~r1_xxreal_0(X0,k2_int_1(sK0)) | ~v1_xxreal_0(k2_int_1(sK0)) | ~v1_xxreal_0(X0) <- (2, 14) [subsumption resolution 309,193] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
312. r1_xxreal_0(X0,k1_int_1(sK0)) | ~r1_xxreal_0(X0,k2_int_1(sK0)) | ~v1_xxreal_0(X0) <- (2, 13, 14) [subsumption resolution 311,189] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
313. 15 | ~2 | ~13 | ~14 [avatar split clause 312,192,188,73,196]
319. ~v1_xreal_0(sK0) <- (~21) [resolution 299,148] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
328. $false <- (3, ~21) [subsumption resolution 319,81] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
329. ~3 | 21 [avatar contradiction clause 328]
359. r1_xxreal_0(X1,k2_int_1(X1)) | ~v1_xreal_0(X1) [subsumption resolution 176,165] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
401. ~v1_xreal_0(sK0) | ~v1_xxreal_0(sK0) | r1_xxreal_0(sK0,k1_int_1(sK0)) <- (15) [resolution 359,197] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
405. ~v1_xxreal_0(sK0) | r1_xxreal_0(sK0,k1_int_1(sK0)) <- (3, 15) [subsumption resolution 401,81] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
406. r1_xxreal_0(sK0,k1_int_1(sK0)) <- (3, 9, 15) [subsumption resolution 405,112] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
408. 28 <=> r1_xxreal_0(sK0,k1_int_1(sK0)) [avatar definition]
410. r1_xxreal_0(sK0,k1_int_1(sK0)) <- (28) [avatar component clause 408] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
411. 28 | ~3 | ~9 | ~15 [avatar split clause 406,196,110,79,408]
493. v1_int_1(sK0) | ~v1_xreal_0(sK0) <- (28) [resolution 410,57] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
499. ~v1_xreal_0(sK0) <- (~1, 28) [subsumption resolution 493,70] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
500. $false <- (~1, 3, 28) [subsumption resolution 499,81] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
501. 1 | ~3 | ~28 [avatar contradiction clause 500]
506. $false [avatar sat refutation 76,77,82,113,203,209,300,313,329,411,501] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:329,thDist:-329}
% SZS output end Proof for int_1__t35_int_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 345
% Active clauses: 64
% Passive clauses: 107
% Final active clauses: 52
% Final passive clauses: 36
% Main loop iterations started: 70
% 
% >>> Simplifying Inferences
% Duplicate literals: 43
% Fw subsumption resolutions: 70
% 
% >>> Deletion Inferences
% Simple tautologies: 21
% Forward subsumptions: 70
% 
% >>> Generating Inferences
% Binary resolution: 93
% Unit resulting resolution: 11
% Factoring: 1
% Forward superposition: 43
% Backward superposition: 29
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 20
% Sat splitting refutations: 21
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 48
% SAT solver unit clauses: 4
% SAT solver binary clauses: 24
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

