i: [1,112,0,2,1,1,0,0] t9_tops_3
a: [3,161,0,2,1,112]
s: 161
i: [1,111,0,2,1,1,0,0] t9_tops_3
a: [3,166,0,2,1,111]
s: 166
i: [1,113,0,4,1,1,0,0] t9_tops_3
a: [3,156,0,4,1,113]
s: 156
i: [1,114,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,171,0,2,1,114]
s: 171
i: [1,151,0,2,1,0,0,2] existence_l1_struct_0
a: [3,181,0,2,1,151]
s: 181
i: [1,152,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,186,0,2,1,152]
s: 186
i: [1,149,0,3,1,0,0,255] reflexivity_r1_tarski
s: 149
i: [1,115,0,3,1,0,0,255] d2_xboole_0
a: [3,176,0,3,1,115]
s: 176
i: [1,119,0,4,2,0,0,2] dt_l1_pre_topc
s: 119
i: [1,148,0,4,1,0,0,2] existence_m1_subset_1
s: 148
i: [1,116,0,5,2,0,0,1] fc3_struct_0
s: 116
d: [2,188,1,2,1,1,40,119,161]
a: [3,193,1,2,1,188]
s: 193
i: [1,123,0,5,2,0,0,255] t6_boole
s: 123
d: [2,189,1,2,1,1,40,119,186]
a: [3,198,1,2,1,189]
s: 198
i: [1,145,0,5,2,0,0,255] t7_boole
s: 145
d: [2,200,0,5,2,1,22,123,176]
s: 200
i: [1,117,0,6,2,0,0,255] d2_struct_0
s: 117
i: [1,122,0,6,2,0,0,1] rc5_tops_1
s: 122
i: [1,141,0,6,2,0,0,255] t1_subset
s: 141
i: [1,150,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 150
d: [2,203,0,6,2,1,22,117,176]
s: 203
d: [2,204,1,5,2,2,44,156,203]
d: [2,207,1,3,1,3,21,204,193]
a: [3,211,1,3,1,207]
s: 211
i: [1,143,0,7,2,0,0,255] t3_subset
s: 143
d: [2,213,2,5,2,2,44,211,200]
s: 213
i: [1,144,0,7,3,0,0,255] t8_boole
s: 144
i: [1,118,0,8,2,0,0,1] dt_k1_struct_0
s: 118
d: [2,214,3,5,2,2,40,213,122]
d: [2,215,3,3,1,3,21,214,161]
a: [3,219,3,3,1,215]
s: 219
i: [1,121,0,8,2,0,0,1] rc5_tops_1
s: 121
d: [2,201,1,7,3,1,44,200,200]
s: 201
i: [1,131,0,8,3,0,0,255] rc6_pre_topc
s: 131
i: [1,133,0,8,3,0,0,2] rc7_tops_1
s: 133
d: [2,221,1,9,3,1,44,118,203]
d: [2,222,1,7,2,1,18,221]
s: 222
i: [1,135,0,8,3,0,0,3] rc10_pre_topc
s: 135
i: [1,137,0,8,3,0,0,3] rc1_tops_1
s: 137
i: [1,139,0,8,3,0,0,255] rc2_tops_1
s: 139
i: [1,140,0,8,3,0,0,255] rc2_tops_1
s: 140
i: [1,142,0,8,3,0,0,255] t2_subset
s: 142
i: [1,147,0,9,3,0,0,255] t5_subset
s: 147
d: [2,231,2,8,3,1,40,147,222]
a: [3,239,2,5,2,231]
s: 239
i: [1,130,0,10,3,0,0,255] rc6_pre_topc
s: 130
i: [1,132,0,10,3,0,0,2] rc7_tops_1
s: 132
d: [2,228,1,6,2,0,40,142,148]
s: 228
i: [1,134,0,10,3,0,0,3] rc10_pre_topc
s: 134
d: [2,245,2,6,2,0,40,228,150]
s: 245
i: [1,136,0,10,3,0,0,3] rc1_tops_1
s: 136
i: [1,138,0,10,3,0,0,255] rc2_tops_1
s: 138
d: [2,233,1,7,2,0,40,147,148]
s: 233
i: [1,146,0,10,3,0,0,255] t4_subset
s: 146
d: [2,254,2,6,2,0,40,233,228]
s: 254
i: [1,120,0,12,4,0,0,1] cc2_tops_1
s: 120
i: [1,124,0,14,5,0,0,255] cc2_pre_topc
s: 124
d: [2,267,2,9,4,1,40,120,222]
d: [2,282,2,7,3,2,21,267,171]
d: [2,283,2,5,2,2,21,282,119]
s: 283
d: [2,303,3,0,0,4,52,161,211,283]
e: 303
d: [2,304,3,2,1,3,40,283,211]
d: [2,308,3,0,0,4,21,304,161]
e: 308
d: [2,305,3,4,2,3,40,283,213]
d: [2,310,3,2,1,4,21,305,161]
d: [2,311,3,0,0,4,21,310,171]
e: 311
f: 311,207,188,115,114,112,113
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tops_3__t9_tops_3
% SZS output start Proof for tops_3__t9_tops_3
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => v2_tops_1(k1_struct_0(X0),X0)) [input t9_tops_3]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => v2_tops_1(k1_struct_0(X0),X0)) [negated conjecture 1]
16. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
19. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v1_xboole_0(X1) => v2_tops_1(X1,X0)))) [input cc2_tops_1]
34. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
35. o_0_0_xboole_0 = k1_xboole_0 [input d2_xboole_0]
37. ! [X0] : (l1_struct_0(X0) => m1_subset_1(k1_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0)))) [input dt_k1_struct_0]
38. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
39. ! [X0] : (l1_struct_0(X0) => k1_struct_0(X0) = k1_xboole_0) [input d2_struct_0]
49. ? [X0] : (~v2_tops_1(k1_struct_0(X0),X0) & (l1_pre_topc(X0) & v2_pre_topc(X0))) [ennf transformation 2]
50. ? [X0] : (~v2_tops_1(k1_struct_0(X0),X0) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 49]
52. ! [X0] : (k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0)) [ennf transformation 39]
53. ! [X0] : (m1_subset_1(k1_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0)) [ennf transformation 37]
54. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 38]
55. ! [X0] : (! [X1] : ((v2_tops_1(X1,X0) | ~v1_xboole_0(X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 19]
56. ! [X0] : (! [X1] : (v2_tops_1(X1,X0) | ~v1_xboole_0(X1) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [flattening 55]
58. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 34]
91. ? [X0] : (~v2_tops_1(k1_struct_0(X0),X0) & l1_pre_topc(X0) & v2_pre_topc(X0)) => (~v2_tops_1(k1_struct_0(sK0),sK0) & l1_pre_topc(sK0) & v2_pre_topc(sK0)) [choice axiom]
92. ~v2_tops_1(k1_struct_0(sK0),sK0) & l1_pre_topc(sK0) & v2_pre_topc(sK0) [skolemisation 50,91]
112. l1_pre_topc(sK0) [cnf transformation 92] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
113. ~v2_tops_1(k1_struct_0(sK0),sK0) [cnf transformation 92] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
114. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 16] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
115. o_0_0_xboole_0 = k1_xboole_0 [cnf transformation 35] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
117. k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0) [cnf transformation 52] {a:0,w:6,wCS:2400,sine:255,thAx:0,allAx:1,thDist:-1}
118. m1_subset_1(k1_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) [cnf transformation 53] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
119. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 54] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
120. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v1_xboole_0(X1) | v2_tops_1(X1,X0) | ~l1_pre_topc(X0) [cnf transformation 56] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
123. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 58] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
154. 1 <=> v2_tops_1(k1_struct_0(sK0),sK0) [avatar definition]
156. ~v2_tops_1(k1_struct_0(sK0),sK0) <- (~1) [avatar component clause 154] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
157. ~1 [avatar split clause 113,154]
159. 2 <=> l1_pre_topc(sK0) [avatar definition]
161. l1_pre_topc(sK0) <- (2) [avatar component clause 159] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
162. 2 [avatar split clause 112,159]
169. 4 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
171. v1_xboole_0(o_0_0_xboole_0) <- (4) [avatar component clause 169] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
172. 4 [avatar split clause 114,169]
174. 5 <=> o_0_0_xboole_0 = k1_xboole_0 [avatar definition]
176. o_0_0_xboole_0 = k1_xboole_0 <- (5) [avatar component clause 174] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
177. 5 [avatar split clause 115,174]
188. l1_struct_0(sK0) <- (2) [resolution 119,161] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
191. 8 <=> l1_struct_0(sK0) [avatar definition]
193. l1_struct_0(sK0) <- (8) [avatar component clause 191] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
194. 8 | ~2 [avatar split clause 188,159,191]
200. o_0_0_xboole_0 = X0 | ~v1_xboole_0(X0) <- (5) [forward demodulation 123,176] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
203. k1_struct_0(X0) = o_0_0_xboole_0 | ~l1_struct_0(X0) <- (5) [forward demodulation 117,176] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
204. ~v2_tops_1(o_0_0_xboole_0,sK0) | ~l1_struct_0(sK0) <- (~1, 5) [superposition 156,203] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
207. ~v2_tops_1(o_0_0_xboole_0,sK0) <- (~1, 5, 8) [subsumption resolution 204,193] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
209. 10 <=> v2_tops_1(o_0_0_xboole_0,sK0) [avatar definition]
211. ~v2_tops_1(o_0_0_xboole_0,sK0) <- (~10) [avatar component clause 209] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
212. ~10 | 1 | ~5 | ~8 [avatar split clause 207,191,174,154,209]
213. ~v2_tops_1(X0,sK0) | ~v1_xboole_0(X0) <- (5, ~10) [superposition 211,200] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
221. m1_subset_1(o_0_0_xboole_0,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) | ~l1_struct_0(X0) <- (5) [superposition 118,203] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
222. m1_subset_1(o_0_0_xboole_0,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) <- (5) [duplicate literal removal 221] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
267. ~v1_xboole_0(o_0_0_xboole_0) | v2_tops_1(o_0_0_xboole_0,X3) | ~l1_pre_topc(X3) | ~l1_struct_0(X3) <- (5) [resolution 120,222] {a:2,w:9,wCS:3600,sine:1,thAx:0,allAx:3,thDist:-3}
282. v2_tops_1(o_0_0_xboole_0,X3) | ~l1_pre_topc(X3) | ~l1_struct_0(X3) <- (4, 5) [subsumption resolution 267,171] {a:2,w:7,wCS:2800,sine:1,thAx:0,allAx:3,thDist:-3}
283. v2_tops_1(o_0_0_xboole_0,X3) | ~l1_pre_topc(X3) <- (4, 5) [subsumption resolution 282,119] {a:2,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
305. ~l1_pre_topc(sK0) | ~v1_xboole_0(o_0_0_xboole_0) <- (4, 5, ~10) [resolution 283,213] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
310. ~v1_xboole_0(o_0_0_xboole_0) <- (2, 4, 5, ~10) [subsumption resolution 305,161] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
311. $false <- (2, 4, 5, ~10) [subsumption resolution 310,171] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
312. ~2 | ~4 | ~5 | 10 [avatar contradiction clause 311]
313. $false [avatar sat refutation 157,162,172,177,194,212,312] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:73,thDist:-73}
% SZS output end Proof for tops_3__t9_tops_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% >>> Preprocessing
% Introduced skolems: 10
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 42
% Generated clauses: 149
% Active clauses: 49
% Passive clauses: 101
% Final active clauses: 49
% Final passive clauses: 49
% Main loop iterations started: 52
% 
% >>> Simplifying Inferences
% Duplicate literals: 19
% Fw subsumption resolutions: 11
% Fw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 59
% Unit resulting resolution: 1
% Forward superposition: 4
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 7
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption index maintenance: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

