i: [1,134,0,2,1,1,0,0] t36_interva1
a: [3,204,0,2,1,134]
s: 204
i: [1,132,0,2,1,1,0,0] t36_interva1
a: [3,214,0,2,1,132]
s: 214
i: [1,131,0,2,1,1,0,0] t36_interva1
a: [3,219,0,2,1,131]
s: 219
i: [1,135,0,3,1,1,0,0] t36_interva1
a: [3,199,0,3,1,135]
s: 199
i: [1,133,0,3,1,1,0,0] t36_interva1
a: [3,209,0,3,1,133]
s: 209
i: [1,137,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,224,0,2,1,137]
s: 224
i: [1,138,0,2,1,0,0,255] fc1_xboole_0
a: [3,229,0,2,1,138]
s: 229
i: [1,187,0,2,1,0,0,1] rc2_xboole_0
a: [3,239,0,2,1,187]
s: 239
i: [1,188,0,2,1,0,0,1] rc1_xboole_0
a: [3,244,0,2,1,188]
s: 244
i: [1,136,0,10,1,1,0,0] t36_interva1
a: [3,194,0,10,1,136]
s: 194
i: [1,140,0,3,1,0,0,2] fc1_subset_1
s: 140
i: [1,169,0,3,1,0,0,2] rc2_subset_1
s: 169
i: [1,173,0,3,1,0,0,255] rc2_interva1
s: 173
i: [1,175,0,3,1,0,0,255] reflexivity_r1_tarski
s: 175
i: [1,139,0,3,1,0,0,255] d2_xboole_0
a: [3,234,0,3,1,139]
s: 234
r: 224
i: [1,166,0,4,1,0,0,1] existence_m1_interva1
s: 166
i: [1,167,0,4,1,0,0,2] existence_m1_subset_1
s: 167
i: [1,171,0,4,1,0,0,255] rc3_subset_1
s: 171
i: [1,174,0,4,1,0,0,255] rc2_interva1
s: 174
i: [1,149,0,5,2,0,0,1] rc1_interva1
s: 149
i: [1,151,0,5,2,0,0,2] rc1_subset_1
s: 151
i: [1,154,0,5,2,0,0,255] t6_boole
s: 154
i: [1,163,0,5,2,0,0,255] t7_boole
s: 163
i: [1,168,0,5,1,0,0,2] rc2_subset_1
s: 168
i: [1,170,0,5,1,0,0,255] rc3_subset_1
s: 170
i: [1,148,0,6,2,0,0,1] rc1_interva1
s: 148
i: [1,153,0,6,2,0,0,255] rc4_subset_1
s: 153
i: [1,157,0,6,2,0,0,255] t1_subset
s: 157
i: [1,172,0,6,1,0,0,255] rc2_interva1
s: 172
i: [1,178,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 178
i: [1,150,0,7,2,0,0,2] rc1_subset_1
s: 150
i: [1,152,0,7,2,0,0,255] rc4_subset_1
s: 152
i: [1,161,0,7,2,0,0,255] t3_subset
s: 161
i: [1,162,0,7,3,0,0,2] t8_boole
s: 162
i: [1,176,0,7,1,0,0,255] commutativity_k2_setfam_1
s: 176
i: [1,177,0,7,1,0,0,255] commutativity_k3_setfam_1
s: 177
d: [2,247,1,7,3,0,44,154,154]
s: 247
i: [1,155,0,8,3,0,0,2] cc1_subset_1
s: 155
d: [2,254,1,5,2,0,40,155,170]
s: 254
i: [1,158,0,8,2,0,0,1] dt_m1_interva1
s: 158
i: [1,159,0,8,3,0,0,255] t2_subset
s: 159
d: [2,262,1,8,2,0,40,159,168]
d: [2,269,1,5,1,0,21,262,140]
s: 269
i: [1,156,0,9,3,0,0,255] cc4_subset_1
s: 156
d: [2,263,1,8,2,0,40,159,170]
d: [2,270,1,5,1,0,21,263,140]
s: 270
i: [1,165,0,9,3,0,0,255] t5_subset
s: 165
i: [1,144,0,10,4,0,0,255] l40_interva1
s: 144
d: [2,294,1,7,3,1,40,144,199]
d: [2,299,1,5,2,2,21,294,204]
d: [2,300,1,3,1,3,21,299,219]
a: [3,304,1,3,1,300]
s: 304
i: [1,164,0,10,3,0,0,255] t4_subset
s: 164
d: [2,295,1,7,3,1,40,144,209]
d: [2,306,1,5,2,2,21,295,214]
d: [2,307,1,3,1,3,21,306,219]
a: [3,311,1,3,1,307]
s: 311
i: [1,141,0,11,4,0,0,255] cc3_subset_1
s: 141
i: [1,142,0,11,4,0,0,255] cc2_subset_1
s: 142
d: [2,328,1,9,3,0,40,141,170]
d: [2,334,1,5,2,0,21,328,171]
s: 334
i: [1,189,0,16,6,0,0,1] redefinition_r1_interva1
d: [2,190,0,11,4,0,18,189]
s: 190
d: [2,272,2,5,1,0,40,269,178]
s: 272
i: [1,145,0,12,4,0,0,255] l40_interva1
s: 145
i: [1,179,0,16,6,0,0,1] reflexivity_r1_interva1
s: 179
d: [2,283,2,5,1,0,40,270,178]
s: 283
i: [1,180,0,17,6,0,0,1] fc1_interva1
s: 180
d: [2,252,1,6,2,0,40,155,167]
s: 252
i: [1,181,0,17,6,0,0,1] fc2_interva1
s: 181
i: [1,182,0,18,6,0,0,1] dt_k4_interva1
s: 182
d: [2,261,1,6,2,0,40,159,167]
s: 261
i: [1,183,0,18,6,0,0,1] dt_k3_interva1
s: 183
d: [2,264,1,10,2,0,40,159,172]
d: [2,271,1,6,1,0,21,264,140]
s: 271
i: [1,185,0,19,7,0,0,1] redefinition_r1_interva1
s: 185
i: [1,146,0,20,6,0,0,255] d4_interva1
s: 146
d: [2,350,1,27,6,1,44,194,146]
d: [2,355,1,24,5,2,21,350,199]
d: [2,356,1,22,4,3,21,355,204]
d: [2,357,1,20,3,4,21,356,219]
d: [2,358,1,20,3,4,22,357,176]
a: [3,366,1,6,1,358]
s: 366
d: [2,372,2,0,0,6,52,219,214,204,209,199,366,183]
e: 372
d: [2,373,2,12,5,1,40,366,183]
d: [2,375,2,9,4,2,21,373,199]
d: [2,376,2,7,3,3,21,375,204]
d: [2,377,2,4,2,4,21,376,209]
d: [2,378,2,2,1,5,21,377,214]
d: [2,379,2,0,0,6,21,378,219]
e: 379
r: 366
i: [1,147,0,20,6,0,0,255] d3_interva1
s: 147
a: [3,365,1,6,1,358]
s: 365
i: [1,184,0,20,7,0,0,1] symmetry_r1_interva1
s: 184
i: [1,160,0,27,7,0,0,255] t34_interva1
s: 160
d: [2,397,2,17,6,1,44,365,147]
d: [2,404,2,14,5,2,21,397,199]
d: [2,405,2,12,4,3,21,404,204]
d: [2,406,2,9,3,4,21,405,209]
d: [2,407,2,7,2,5,21,406,214]
d: [2,408,2,5,1,6,21,407,219]
a: [3,412,2,5,1,408]
s: 412
d: [2,416,0,27,7,0,22,160,176]
s: 416
d: [2,417,3,11,3,1,40,412,144]
d: [2,418,3,9,2,2,21,417,219]
a: [3,422,3,4,1,418]
s: 422
a: [3,370,1,9,1,358]
s: 370
d: [2,381,1,21,6,1,44,194,147]
d: [2,386,1,18,5,2,21,381,199]
d: [2,387,1,16,4,3,21,386,204]
d: [2,388,1,13,3,4,21,387,209]
d: [2,389,1,11,2,5,21,388,214]
d: [2,390,1,9,1,6,21,389,219]
a: [3,394,1,9,1,390]
s: 394
d: [2,396,2,13,3,1,40,365,144]
d: [2,398,2,11,2,2,21,396,219]
a: [3,402,2,6,1,398]
s: 402
d: [2,280,1,6,2,0,40,156,168]
s: 280
a: [3,426,3,5,1,418]
s: 426
d: [2,291,1,6,2,0,40,165,168]
s: 291
d: [2,292,1,6,2,0,40,165,170]
s: 292
d: [2,430,2,20,6,1,44,370,147]
d: [2,431,2,17,5,2,21,430,199]
d: [2,432,2,15,4,3,21,431,204]
d: [2,433,2,12,3,4,21,432,209]
d: [2,434,2,10,2,5,21,433,214]
d: [2,435,2,8,1,6,21,434,219]
a: [3,439,2,8,1,435]
s: 439
d: [2,297,1,11,4,0,40,144,148]
d: [2,298,1,9,3,0,18,297]
d: [2,313,1,6,2,0,21,298,149]
s: 313
d: [2,451,3,24,7,1,44,439,416]
d: [2,452,3,22,6,2,21,451,204]
d: [2,453,3,20,5,3,21,452,214]
a: [3,460,3,4,1,453]
s: 460
d: [2,462,4,0,0,4,52,219,204,199,460,190]
e: 462
d: [2,463,4,7,3,1,40,460,190]
d: [2,465,4,4,2,2,21,463,199]
d: [2,466,4,2,1,3,21,465,204]
d: [2,467,4,0,0,4,21,466,219]
e: 467
r: 460
d: [2,327,1,9,3,0,40,141,168]
d: [2,333,1,6,2,0,21,327,169]
s: 333
d: [2,249,1,7,3,0,40,155,161]
s: 249
a: [3,459,3,4,1,453]
s: 459
d: [2,257,1,10,3,0,40,159,161]
d: [2,265,1,7,2,0,21,257,140]
s: 265
a: [3,456,3,16,4,453]
s: 456
d: [2,259,1,10,3,0,40,159,150]
d: [2,267,1,7,2,0,21,259,140]
s: 267
d: [2,260,1,10,3,0,40,159,152]
d: [2,268,1,7,2,0,21,260,140]
s: 268
d: [2,478,4,14,4,1,40,456,158]
s: 478
d: [2,279,1,7,2,0,40,156,167]
s: 279
d: [2,486,5,12,4,1,40,478,158]
s: 486
d: [2,488,6,0,0,5,52,209,199,311,304,486]
e: 488
d: [2,489,6,9,3,2,40,486,304]
d: [2,491,6,6,2,3,21,489,311]
d: [2,492,6,3,1,4,21,491,199]
d: [2,493,6,0,0,5,21,492,209]
e: 493
r: 370
r: 439
r: 486
r: 478
r: 456
a: [3,362,1,5,1,358]
s: 362
d: [2,495,2,0,0,6,52,219,214,204,209,199,362,180]
e: 495
d: [2,496,2,12,5,1,40,362,180]
d: [2,499,2,9,4,2,21,496,199]
d: [2,500,2,7,3,3,21,499,204]
d: [2,501,2,4,2,4,21,500,209]
d: [2,502,2,2,1,5,21,501,214]
d: [2,503,2,0,0,6,21,502,219]
e: 503
f: 503,493,467,453,435,379,358,307,300,131,132,133,134,135,136
% Refutation found. Thanks to Tanya!
% SZS status Theorem for interva1__t36_interva1
% SZS output start Proof for interva1__t36_interva1
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2)) => r1_interva1(X0,k4_interva1(X0,k3_interva1(X0,X1,X2),X2),X2)))) [input t36_interva1]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2)) => r1_interva1(X0,k4_interva1(X0,k3_interva1(X0,X1,X2),X2),X2)))) [negated conjecture 1]
8. ! [X0,X1] : k2_setfam_1(X0,X1) = k2_setfam_1(X1,X0) [input commutativity_k2_setfam_1]
11. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2)) => k3_interva1(X0,X1,X2) = k3_setfam_1(X1,X2)))) [input d3_interva1]
12. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2)) => k4_interva1(X0,X1,X2) = k2_setfam_1(X1,X2)))) [input d4_interva1]
16. ! [X0,X1,X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2) & m1_interva1(X1,X0) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => m1_interva1(k3_interva1(X0,X1,X2),X0)) [input dt_k3_interva1]
19. ! [X0] : ! [X1] : (m1_interva1(X1,X0) => m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [input dt_m1_interva1]
24. ! [X0,X1,X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2) & m1_interva1(X1,X0) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k3_interva1(X0,X1,X2))) [input fc1_interva1]
28. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)))) [input l40_interva1]
37. ! [X0,X1,X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2) & m1_interva1(X1,X0) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => (r1_interva1(X0,X1,X2) <=> X1 = X2)) [input redefinition_r1_interva1]
43. ! [X0] : ! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X2,X0) & ~v1_xboole_0(X2)) => k2_setfam_1(k3_setfam_1(X1,X2),X2) = X2)) [input t34_interva1]
51. ! [X0,X1] : (m1_interva1(X1,X0) => m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [flattening 19]
53. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X2,X0) & ~v1_xboole_0(X2)) => k2_setfam_1(k3_setfam_1(X1,X2),X2) = X2)) [flattening 43]
61. ? [X0] : (? [X1] : (? [X2] : (~r1_interva1(X0,k4_interva1(X0,k3_interva1(X0,X1,X2),X2),X2) & (m1_interva1(X2,X0) & ~v1_xboole_0(X2))) & (m1_interva1(X1,X0) & ~v1_xboole_0(X1))) & ~v1_xboole_0(X0)) [ennf transformation 2]
62. ? [X0] : (? [X1] : (? [X2] : (~r1_interva1(X0,k4_interva1(X0,k3_interva1(X0,X1,X2),X2),X2) & m1_interva1(X2,X0) & ~v1_xboole_0(X2)) & m1_interva1(X1,X0) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [flattening 61]
67. ! [X0] : (! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)) | (~m1_interva1(X1,X0) | v1_xboole_0(X1))) | v1_xboole_0(X0)) [ennf transformation 28]
68. ! [X0] : (! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)) | ~m1_interva1(X1,X0) | v1_xboole_0(X1)) | v1_xboole_0(X0)) [flattening 67]
69. ! [X0] : (! [X1] : (! [X2] : (k4_interva1(X0,X1,X2) = k2_setfam_1(X1,X2) | (~m1_interva1(X2,X0) | v1_xboole_0(X2))) | (~m1_interva1(X1,X0) | v1_xboole_0(X1))) | v1_xboole_0(X0)) [ennf transformation 12]
70. ! [X0] : (! [X1] : (! [X2] : (k4_interva1(X0,X1,X2) = k2_setfam_1(X1,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2)) | ~m1_interva1(X1,X0) | v1_xboole_0(X1)) | v1_xboole_0(X0)) [flattening 69]
71. ! [X0] : (! [X1] : (! [X2] : (k3_interva1(X0,X1,X2) = k3_setfam_1(X1,X2) | (~m1_interva1(X2,X0) | v1_xboole_0(X2))) | (~m1_interva1(X1,X0) | v1_xboole_0(X1))) | v1_xboole_0(X0)) [ennf transformation 11]
72. ! [X0] : (! [X1] : (! [X2] : (k3_interva1(X0,X1,X2) = k3_setfam_1(X1,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2)) | ~m1_interva1(X1,X0) | v1_xboole_0(X1)) | v1_xboole_0(X0)) [flattening 71]
80. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~m1_interva1(X1,X0)) [ennf transformation 51]
83. ! [X0,X1] : (! [X2] : (k2_setfam_1(k3_setfam_1(X1,X2),X2) = X2 | (~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X2,X0) | v1_xboole_0(X2))) | (~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X1,X0) | v1_xboole_0(X1))) [ennf transformation 53]
84. ! [X0,X1] : (! [X2] : (k2_setfam_1(k3_setfam_1(X1,X2),X2) = X2 | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X2,X0) | v1_xboole_0(X2)) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X1,X0) | v1_xboole_0(X1)) [flattening 83]
94. ! [X0,X1,X2] : (~v1_xboole_0(k3_interva1(X0,X1,X2)) | (~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 24]
95. ! [X0,X1,X2] : (~v1_xboole_0(k3_interva1(X0,X1,X2)) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 94]
100. ! [X0,X1,X2] : (m1_interva1(k3_interva1(X0,X1,X2),X0) | (~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 16]
101. ! [X0,X1,X2] : (m1_interva1(k3_interva1(X0,X1,X2),X0) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 100]
104. ! [X0,X1,X2] : ((r1_interva1(X0,X1,X2) <=> X1 = X2) | (~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 37]
105. ! [X0,X1,X2] : ((r1_interva1(X0,X1,X2) <=> X1 = X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 104]
106. ? [X0] : (? [X1] : (? [X2] : (~r1_interva1(X0,k4_interva1(X0,k3_interva1(X0,X1,X2),X2),X2) & m1_interva1(X2,X0) & ~v1_xboole_0(X2)) & m1_interva1(X1,X0) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : (~r1_interva1(sK0,k4_interva1(sK0,k3_interva1(sK0,X1,X2),X2),X2) & m1_interva1(X2,sK0) & ~v1_xboole_0(X2)) & m1_interva1(X1,sK0) & ~v1_xboole_0(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
107. ? [X1] : (? [X2] : (~r1_interva1(sK0,k4_interva1(sK0,k3_interva1(sK0,X1,X2),X2),X2) & m1_interva1(X2,sK0) & ~v1_xboole_0(X2)) & m1_interva1(X1,sK0) & ~v1_xboole_0(X1)) => (? [X2] : (~r1_interva1(sK0,k4_interva1(sK0,k3_interva1(sK0,sK1,X2),X2),X2) & m1_interva1(X2,sK0) & ~v1_xboole_0(X2)) & m1_interva1(sK1,sK0) & ~v1_xboole_0(sK1)) [choice axiom]
108. ? [X2] : (~r1_interva1(sK0,k4_interva1(sK0,k3_interva1(sK0,sK1,X2),X2),X2) & m1_interva1(X2,sK0) & ~v1_xboole_0(X2)) => (~r1_interva1(sK0,k4_interva1(sK0,k3_interva1(sK0,sK1,sK2),sK2),sK2) & m1_interva1(sK2,sK0) & ~v1_xboole_0(sK2)) [choice axiom]
109. ((~r1_interva1(sK0,k4_interva1(sK0,k3_interva1(sK0,sK1,sK2),sK2),sK2) & m1_interva1(sK2,sK0) & ~v1_xboole_0(sK2)) & m1_interva1(sK1,sK0) & ~v1_xboole_0(sK1)) & ~v1_xboole_0(sK0) [skolemisation 62,108,107,106]
126. ! [X0,X1,X2] : (((r1_interva1(X0,X1,X2) | X1 != X2) & (X1 = X2 | ~r1_interva1(X0,X1,X2))) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0)) [nnf transformation 105]
131. ~v1_xboole_0(sK0) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
132. ~v1_xboole_0(sK1) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
133. m1_interva1(sK1,sK0) [cnf transformation 109] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
134. ~v1_xboole_0(sK2) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
135. m1_interva1(sK2,sK0) [cnf transformation 109] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
136. ~r1_interva1(sK0,k4_interva1(sK0,k3_interva1(sK0,sK1,sK2),sK2),sK2) [cnf transformation 109] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
144. ~m1_interva1(X1,X0) | v1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 68] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
146. k4_interva1(X0,X1,X2) = k2_setfam_1(X1,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 70] {a:0,w:20,wCS:8000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
147. k3_interva1(X0,X1,X2) = k3_setfam_1(X1,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 72] {a:0,w:20,wCS:8000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
158. m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~m1_interva1(X1,X0) [cnf transformation 80] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
160. k2_setfam_1(k3_setfam_1(X1,X2),X2) = X2 | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X1,X0) | v1_xboole_0(X1) [cnf transformation 84] {a:0,w:27,wCS:10800,sine:255,thAx:0,allAx:1,thDist:-1}
176. k2_setfam_1(X0,X1) = k2_setfam_1(X1,X0) [cnf transformation 8] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
180. ~v1_xboole_0(k3_interva1(X0,X1,X2)) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 95] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
183. m1_interva1(k3_interva1(X0,X1,X2),X0) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 101] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
186. r1_interva1(X0,X1,X2) | X1 != X2 | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 126] {a:0,w:19,wCS:7600,sine:1,thAx:0,allAx:1,thDist:-1}
189. r1_interva1(X0,X2,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | v1_xboole_0(X0) [equality resolution 186] {a:0,w:16,wCS:6400,sine:1,thAx:0,allAx:1,thDist:-1}
190. r1_interva1(X0,X2,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | v1_xboole_0(X0) [duplicate literal removal 189] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
192. 1 <=> r1_interva1(sK0,k4_interva1(sK0,k3_interva1(sK0,sK1,sK2),sK2),sK2) [avatar definition]
194. ~r1_interva1(sK0,k4_interva1(sK0,k3_interva1(sK0,sK1,sK2),sK2),sK2) <- (~1) [avatar component clause 192] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
195. ~1 [avatar split clause 136,192]
197. 2 <=> m1_interva1(sK2,sK0) [avatar definition]
199. m1_interva1(sK2,sK0) <- (2) [avatar component clause 197] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
200. 2 [avatar split clause 135,197]
202. 3 <=> v1_xboole_0(sK2) [avatar definition]
204. ~v1_xboole_0(sK2) <- (~3) [avatar component clause 202] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
205. ~3 [avatar split clause 134,202]
207. 4 <=> m1_interva1(sK1,sK0) [avatar definition]
209. m1_interva1(sK1,sK0) <- (4) [avatar component clause 207] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
210. 4 [avatar split clause 133,207]
212. 5 <=> v1_xboole_0(sK1) [avatar definition]
214. ~v1_xboole_0(sK1) <- (~5) [avatar component clause 212] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
215. ~5 [avatar split clause 132,212]
217. 6 <=> v1_xboole_0(sK0) [avatar definition]
219. ~v1_xboole_0(sK0) <- (~6) [avatar component clause 217] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
220. ~6 [avatar split clause 131,217]
294. v1_interva1(sK2,sK0) | v1_xboole_0(sK2) | v1_xboole_0(sK0) <- (2) [resolution 144,199] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
295. v1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (4) [resolution 144,209] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
299. v1_interva1(sK2,sK0) | v1_xboole_0(sK0) <- (2, ~3) [subsumption resolution 294,204] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
300. v1_interva1(sK2,sK0) <- (2, ~3, ~6) [subsumption resolution 299,219] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
302. 12 <=> v1_interva1(sK2,sK0) [avatar definition]
304. v1_interva1(sK2,sK0) <- (12) [avatar component clause 302] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
305. 12 | ~2 | 3 | 6 [avatar split clause 300,217,202,197,302]
306. v1_interva1(sK1,sK0) | v1_xboole_0(sK0) <- (4, ~5) [subsumption resolution 295,214] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
307. v1_interva1(sK1,sK0) <- (4, ~5, ~6) [subsumption resolution 306,219] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
309. 13 <=> v1_interva1(sK1,sK0) [avatar definition]
311. v1_interva1(sK1,sK0) <- (13) [avatar component clause 309] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
312. 13 | ~4 | 5 | 6 [avatar split clause 307,217,212,207,309]
350. ~r1_interva1(sK0,k2_setfam_1(k3_interva1(sK0,sK1,sK2),sK2),sK2) | ~m1_interva1(sK2,sK0) | v1_xboole_0(sK2) | ~m1_interva1(k3_interva1(sK0,sK1,sK2),sK0) | v1_xboole_0(k3_interva1(sK0,sK1,sK2)) | v1_xboole_0(sK0) <- (~1) [superposition 194,146] {a:1,w:27,wCS:2700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
355. ~r1_interva1(sK0,k2_setfam_1(k3_interva1(sK0,sK1,sK2),sK2),sK2) | v1_xboole_0(sK2) | ~m1_interva1(k3_interva1(sK0,sK1,sK2),sK0) | v1_xboole_0(k3_interva1(sK0,sK1,sK2)) | v1_xboole_0(sK0) <- (~1, 2) [subsumption resolution 350,199] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
356. ~r1_interva1(sK0,k2_setfam_1(k3_interva1(sK0,sK1,sK2),sK2),sK2) | ~m1_interva1(k3_interva1(sK0,sK1,sK2),sK0) | v1_xboole_0(k3_interva1(sK0,sK1,sK2)) | v1_xboole_0(sK0) <- (~1, 2, ~3) [subsumption resolution 355,204] {a:1,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
357. ~r1_interva1(sK0,k2_setfam_1(k3_interva1(sK0,sK1,sK2),sK2),sK2) | ~m1_interva1(k3_interva1(sK0,sK1,sK2),sK0) | v1_xboole_0(k3_interva1(sK0,sK1,sK2)) <- (~1, 2, ~3, ~6) [subsumption resolution 356,219] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
358. ~r1_interva1(sK0,k2_setfam_1(sK2,k3_interva1(sK0,sK1,sK2)),sK2) | ~m1_interva1(k3_interva1(sK0,sK1,sK2),sK0) | v1_xboole_0(k3_interva1(sK0,sK1,sK2)) <- (~1, 2, ~3, ~6) [forward demodulation 357,176] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
360. 14 <=> v1_xboole_0(k3_interva1(sK0,sK1,sK2)) [avatar definition]
362. v1_xboole_0(k3_interva1(sK0,sK1,sK2)) <- (14) [avatar component clause 360] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
364. 15 <=> m1_interva1(k3_interva1(sK0,sK1,sK2),sK0) [avatar definition]
366. ~m1_interva1(k3_interva1(sK0,sK1,sK2),sK0) <- (~15) [avatar component clause 364] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
368. 16 <=> r1_interva1(sK0,k2_setfam_1(sK2,k3_interva1(sK0,sK1,sK2)),sK2) [avatar definition]
370. ~r1_interva1(sK0,k2_setfam_1(sK2,k3_interva1(sK0,sK1,sK2)),sK2) <- (~16) [avatar component clause 368] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
371. 14 | ~15 | ~16 | 1 | ~2 | 3 | 6 [avatar split clause 358,217,202,197,192,368,364,360]
373. ~m1_interva1(sK2,sK0) | v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~15) [resolution 366,183] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
375. v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~15) [subsumption resolution 373,199] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
376. ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, ~15) [subsumption resolution 375,204] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
377. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, 4, ~15) [subsumption resolution 376,209] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
378. v1_xboole_0(sK0) <- (2, ~3, 4, ~5, ~15) [subsumption resolution 377,214] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
379. $false <- (2, ~3, 4, ~5, ~6, ~15) [subsumption resolution 378,219] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
380. ~2 | 3 | ~4 | 5 | 6 | 15 [avatar contradiction clause 379]
416. k2_setfam_1(X2,k3_setfam_1(X1,X2)) = X2 | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X1,X0) | v1_xboole_0(X1) [forward demodulation 160,176] {a:0,w:27,wCS:10800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
430. ~r1_interva1(sK0,k2_setfam_1(sK2,k3_setfam_1(sK1,sK2)),sK2) | ~m1_interva1(sK2,sK0) | v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~16) [superposition 370,147] {a:2,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
431. ~r1_interva1(sK0,k2_setfam_1(sK2,k3_setfam_1(sK1,sK2)),sK2) | v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~16) [subsumption resolution 430,199] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
432. ~r1_interva1(sK0,k2_setfam_1(sK2,k3_setfam_1(sK1,sK2)),sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, ~16) [subsumption resolution 431,204] {a:2,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
433. ~r1_interva1(sK0,k2_setfam_1(sK2,k3_setfam_1(sK1,sK2)),sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, 4, ~16) [subsumption resolution 432,209] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
434. ~r1_interva1(sK0,k2_setfam_1(sK2,k3_setfam_1(sK1,sK2)),sK2) | v1_xboole_0(sK0) <- (2, ~3, 4, ~5, ~16) [subsumption resolution 433,214] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
435. ~r1_interva1(sK0,k2_setfam_1(sK2,k3_setfam_1(sK1,sK2)),sK2) <- (2, ~3, 4, ~5, ~6, ~16) [subsumption resolution 434,219] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
437. 22 <=> r1_interva1(sK0,k2_setfam_1(sK2,k3_setfam_1(sK1,sK2)),sK2) [avatar definition]
439. ~r1_interva1(sK0,k2_setfam_1(sK2,k3_setfam_1(sK1,sK2)),sK2) <- (~22) [avatar component clause 437] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
440. ~22 | ~2 | 3 | ~4 | 5 | 6 | 16 [avatar split clause 435,368,217,212,207,202,197,437]
451. ~r1_interva1(sK0,sK2,sK2) | ~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0) | v1_xboole_0(sK2) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | v1_xboole_0(sK1) <- (~22) [superposition 439,416] {a:3,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
452. ~r1_interva1(sK0,sK2,sK2) | ~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | v1_xboole_0(sK1) <- (~3, ~22) [subsumption resolution 451,204] {a:3,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
453. ~r1_interva1(sK0,sK2,sK2) | ~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) <- (~3, ~5, ~22) [subsumption resolution 452,214] {a:3,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
455. 23 <=> ! [X0] : (~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0)) [avatar definition]
456. ~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0) <- (23) [avatar component clause 455] {a:3,w:16,wCS:1600,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
458. 24 <=> r1_interva1(sK0,sK2,sK2) [avatar definition]
460. ~r1_interva1(sK0,sK2,sK2) <- (~24) [avatar component clause 458] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
461. 23 | ~24 | 3 | 5 | 22 [avatar split clause 453,437,212,202,458,455]
463. ~m1_interva1(sK2,sK0) | v1_xboole_0(sK2) | v1_xboole_0(sK0) <- (~24) [resolution 460,190] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
465. v1_xboole_0(sK2) | v1_xboole_0(sK0) <- (2, ~24) [subsumption resolution 463,199] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
466. v1_xboole_0(sK0) <- (2, ~3, ~24) [subsumption resolution 465,204] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
467. $false <- (2, ~3, ~6, ~24) [subsumption resolution 466,219] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
468. ~2 | 3 | 6 | 24 [avatar contradiction clause 467]
478. ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | ~v1_interva1(sK2,X0) | ~m1_interva1(sK2,X0) <- (23) [resolution 456,158] {a:4,w:14,wCS:1400,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
486. ~v1_interva1(sK2,X0) | ~v1_interva1(sK1,X0) | ~m1_interva1(sK2,X0) | ~m1_interva1(sK1,X0) <- (23) [resolution 478,158] {a:5,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
489. ~v1_interva1(sK1,sK0) | ~m1_interva1(sK2,sK0) | ~m1_interva1(sK1,sK0) <- (12, 23) [resolution 486,304] {a:6,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
491. ~m1_interva1(sK2,sK0) | ~m1_interva1(sK1,sK0) <- (12, 13, 23) [subsumption resolution 489,311] {a:6,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
492. ~m1_interva1(sK1,sK0) <- (2, 12, 13, 23) [subsumption resolution 491,199] {a:6,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
493. $false <- (2, 4, 12, 13, 23) [subsumption resolution 492,209] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
494. ~2 | ~4 | ~12 | ~13 | ~23 [avatar contradiction clause 493]
496. ~m1_interva1(sK2,sK0) | v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (14) [resolution 362,180] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
499. v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, 14) [subsumption resolution 496,199] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
500. ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, 14) [subsumption resolution 499,204] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
501. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, 4, 14) [subsumption resolution 500,209] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
502. v1_xboole_0(sK0) <- (2, ~3, 4, ~5, 14) [subsumption resolution 501,214] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
503. $false <- (2, ~3, 4, ~5, ~6, 14) [subsumption resolution 502,219] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
504. ~2 | 3 | ~4 | 5 | 6 | ~14 [avatar contradiction clause 503]
505. $false [avatar sat refutation 195,200,205,210,215,220,305,312,371,380,440,461,468,494,504] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:290,thDist:-290}
% SZS output end Proof for interva1__t36_interva1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 48
% >>> Preprocessing
% Introduced skolems: 13
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 58
% Generated clauses: 277
% Active clauses: 90
% Passive clauses: 128
% Final active clauses: 82
% Final passive clauses: 31
% Main loop iterations started: 95
% 
% >>> Simplifying Inferences
% Duplicate literals: 41
% Fw subsumption resolutions: 60
% Fw demodulations: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 12
% Equational tautologies: 2
% Forward subsumptions: 35
% 
% >>> Generating Inferences
% Binary resolution: 94
% Unit resulting resolution: 4
% Forward superposition: 8
% Backward superposition: 7
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 9
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 31
% SAT solver unit clauses: 11
% 
% Memory used [KB]: 9722
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% term sharing: 0.001 s
% resolution: 0.001 s
% unit resulting resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.004 s ( own 0.001 s ) 

