i: [1,48,0,3,1,1,0,0] t92_zfmisc_1
a: [3,71,0,3,1,48]
s: 71
i: [1,47,0,7,1,1,0,0] t92_zfmisc_1
a: [3,76,0,7,1,47]
s: 76
i: [1,49,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,81,0,2,1,49]
s: 81
i: [1,50,0,2,1,0,0,255] fc1_xboole_0
a: [3,86,0,2,1,50]
s: 86
i: [1,66,0,2,1,0,0,2] rc2_xboole_0
a: [3,96,0,2,1,66]
s: 96
i: [1,67,0,2,1,0,0,2] rc1_xboole_0
a: [3,101,0,2,1,67]
s: 101
i: [1,52,0,3,1,0,0,255] fc2_xboole_0
s: 52
i: [1,51,0,3,1,0,0,255] d2_xboole_0
a: [3,91,0,3,1,51]
s: 91
r: 81
i: [1,62,0,4,1,0,0,2] fc1_zfmisc_1
s: 62
i: [1,63,0,4,1,0,0,255] fc3_xboole_0
s: 63
i: [1,53,0,5,2,0,0,255] t6_boole
s: 53
i: [1,58,0,5,2,0,0,2] t7_boole
s: 58
i: [1,65,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 65
i: [1,57,0,7,3,0,0,1] t8_boole
s: 57
i: [1,64,0,7,1,0,0,255] commutativity_k2_tarski
s: 64
i: [1,54,0,10,1,0,0,255] d5_tarski
s: 54
d: [2,104,1,7,3,0,44,53,53]
s: 104
i: [1,59,0,10,2,0,0,1] l46_zfmisc_1
s: 59
d: [2,116,1,10,2,1,44,59,76]
s: 116
i: [1,60,0,10,2,0,0,1] l46_zfmisc_1
s: 60
i: [1,55,0,13,3,0,0,2] t1_tarski
s: 55
d: [2,117,1,10,2,1,44,60,76]
s: 117
i: [1,56,0,13,3,0,0,2] t1_tarski
s: 56
d: [2,108,1,10,1,0,44,54,64]
s: 108
i: [1,61,0,13,3,0,0,1] l46_zfmisc_1
s: 61
d: [2,139,1,13,3,1,44,61,76]
s: 139
d: [2,150,2,9,3,1,40,139,60]
a: [3,156,2,3,1,150]
s: 156
r: 139
d: [2,109,1,10,1,0,44,54,64]
s: 109
d: [2,135,2,9,3,1,40,61,117]
a: [3,145,2,6,2,135]
s: 145
r: 156
s: 139
d: [2,111,1,10,1,0,44,54,64]
s: 111
d: [2,153,2,10,3,1,40,139,58]
a: [3,165,2,4,1,153]
s: 165
d: [2,112,1,10,1,0,44,54,64]
s: 112
s: 145
d: [2,113,1,10,1,0,44,64,54]
s: 113
d: [2,114,1,10,1,0,44,64,54]
s: 114
a: [3,159,2,6,2,150]
s: 159
d: [2,119,1,10,3,0,40,55,58]
s: 119
d: [2,192,3,10,3,1,40,119,145]
s: 192
d: [2,138,1,10,3,0,40,61,58]
s: 138
d: [2,118,1,13,3,0,40,55,65]
s: 118
d: [2,195,4,13,4,1,40,192,56]
d: [2,199,4,10,3,1,18,195]
d: [2,200,4,7,2,2,21,199,71]
a: [3,208,4,2,1,200]
s: 208
d: [2,137,1,13,3,0,40,61,65]
s: 137
d: [2,193,3,10,3,1,40,119,159]
s: 193
d: [2,110,1,15,1,0,44,54,54]
s: 110
d: [2,152,2,13,3,1,40,139,65]
s: 152
d: [2,216,4,13,4,1,40,193,56]
d: [2,220,4,10,3,1,18,216]
d: [2,221,4,7,2,2,21,220,71]
a: [3,229,4,2,1,221]
s: 229
d: [2,174,2,13,3,1,40,139,65]
s: 174
d: [2,197,4,10,3,1,40,192,65]
s: 197
d: [2,215,2,13,3,1,44,137,76]
s: 215
d: [2,127,2,10,1,0,44,108,64]
s: 127
d: [2,218,4,10,3,1,40,193,65]
s: 218
d: [2,128,2,10,1,0,44,108,64]
s: 128
d: [2,187,3,13,3,1,40,145,55]
s: 187
d: [2,130,2,10,1,0,44,64,108]
s: 130
d: [2,131,2,10,1,0,44,64,108]
s: 131
a: [3,224,4,5,1,221]
s: 224
d: [2,179,2,10,1,0,44,111,64]
s: 179
d: [2,263,5,5,1,2,40,224,159]
a: [3,269,5,5,1,263]
s: 269
d: [2,276,6,0,0,3,52,71,224,269,56]
e: 276
d: [2,277,6,8,2,1,40,269,56]
d: [2,282,6,5,1,2,21,277,71]
d: [2,283,6,0,0,3,21,282,224]
e: 283
r: 218
r: 193
r: 159
r: 269
a: [3,142,2,3,1,135]
s: 142
r: 145
s: 156
r: 139
r: 152
d: [2,297,5,0,0,2,40,156,224]
e: 297
d: [2,288,5,0,0,2,52,224,156]
e: 288
d: [2,289,3,0,0,3,52,71,142,156,55]
e: 289
d: [2,290,3,0,0,3,52,71,142,156,55]
e: 290
d: [2,291,4,0,0,3,52,71,156,156,187]
e: 291
d: [2,292,4,0,0,3,52,71,156,156,187]
e: 292
r: 197
r: 192
r: 187
r: 156
r: 224
d: [2,136,2,9,3,1,40,61,116]
d: [2,173,2,9,3,1,40,139,59]
d: [2,260,4,8,2,1,42,187]
d: [2,261,4,5,1,2,21,260,71]
d: [2,304,2,3,1,2,21,159,142]
f: 304,292,290,283,263,261,173,136,47,48
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t92_zfmisc_1
% SZS output start Proof for zfmisc_1__t92_zfmisc_1
1. ! [X0] : ! [X1] : (k2_zfmisc_1(X0,X0) = k2_zfmisc_1(X1,X1) => X0 = X1) [input t92_zfmisc_1]
2. ~! [X0] : ! [X1] : (k2_zfmisc_1(X0,X0) = k2_zfmisc_1(X1,X1) => X0 = X1) [negated conjecture 1]
17. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [input l46_zfmisc_1]
20. ! [X0] : ! [X1] : (! [X2] : (r2_hidden(X2,X0) <=> r2_hidden(X2,X1)) => X0 = X1) [input t1_tarski]
24. ~! [X0,X1] : (k2_zfmisc_1(X0,X0) = k2_zfmisc_1(X1,X1) => X0 = X1) [flattening 2]
26. ! [X0,X1] : (! [X2] : (r2_hidden(X2,X0) <=> r2_hidden(X2,X1)) => X0 = X1) [flattening 20]
29. ! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [flattening 17]
30. ? [X0,X1] : (X0 != X1 & k2_zfmisc_1(X0,X0) = k2_zfmisc_1(X1,X1)) [ennf transformation 24]
32. ! [X0,X1] : (X0 = X1 | ? [X2] : (r2_hidden(X2,X0) <~> r2_hidden(X2,X1))) [ennf transformation 26]
36. ? [X0,X1] : (X0 != X1 & k2_zfmisc_1(X0,X0) = k2_zfmisc_1(X1,X1)) => (sK0 != sK1 & k2_zfmisc_1(sK0,sK0) = k2_zfmisc_1(sK1,sK1)) [choice axiom]
37. sK0 != sK1 & k2_zfmisc_1(sK0,sK0) = k2_zfmisc_1(sK1,sK1) [skolemisation 30,36]
38. ! [X0,X1] : (X0 = X1 | ? [X2] : ((~r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) & (r2_hidden(X2,X1) | r2_hidden(X2,X0)))) [nnf transformation 32]
39. ! [X1,X0] : (? [X2] : ((~r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) & (r2_hidden(X2,X1) | r2_hidden(X2,X0))) => ((~r2_hidden(sK2(X0,X1),X1) | ~r2_hidden(sK2(X0,X1),X0)) & (r2_hidden(sK2(X0,X1),X1) | r2_hidden(sK2(X0,X1),X0)))) [choice axiom]
40. ! [X0,X1] : (X0 = X1 | ((~r2_hidden(sK2(X0,X1),X1) | ~r2_hidden(sK2(X0,X1),X0)) & (r2_hidden(sK2(X0,X1),X1) | r2_hidden(sK2(X0,X1),X0)))) [skolemisation 38,39]
41. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | (~r2_hidden(X1,X3) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [nnf transformation 29]
42. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [flattening 41]
47. k2_zfmisc_1(sK0,sK0) = k2_zfmisc_1(sK1,sK1) [cnf transformation 37] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. sK0 != sK1 [cnf transformation 37] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. r2_hidden(sK2(X0,X1),X1) | X0 = X1 | r2_hidden(sK2(X0,X1),X0) [cnf transformation 40] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
56. ~r2_hidden(sK2(X0,X1),X1) | X0 = X1 | ~r2_hidden(sK2(X0,X1),X0) [cnf transformation 40] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
59. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X0,X2) [cnf transformation 42] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2) [cnf transformation 42] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. 1 <=> sK0 = sK1 [avatar definition]
71. sK0 != sK1 <- (~1) [avatar component clause 69] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. ~1 [avatar split clause 48,69]
74. 2 <=> k2_zfmisc_1(sK0,sK0) = k2_zfmisc_1(sK1,sK1) [avatar definition]
76. k2_zfmisc_1(sK0,sK0) = k2_zfmisc_1(sK1,sK1) <- (2) [avatar component clause 74] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
77. 2 [avatar split clause 47,74]
116. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(sK0,sK0)) | r2_hidden(X0,sK1) <- (2) [superposition 59,76] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
136. ~r2_hidden(X10,sK0) | ~r2_hidden(X11,sK0) | r2_hidden(X11,sK1) <- (2) [resolution 61,116] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
139. r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(sK0,sK0)) | ~r2_hidden(X1,sK1) | ~r2_hidden(X0,sK1) <- (2) [superposition 61,76] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
141. 8 <=> ! [X9] : ~r2_hidden(X9,sK0) [avatar definition]
142. ~r2_hidden(X9,sK0) <- (8) [avatar component clause 141] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
144. 9 <=> ! [X8] : (~r2_hidden(X8,sK0) | r2_hidden(X8,sK1)) [avatar definition]
145. ~r2_hidden(X8,sK0) | r2_hidden(X8,sK1) <- (9) [avatar component clause 144] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
147. 9 | 8 | ~2 [avatar split clause 136,74,141,144]
155. 10 <=> ! [X5] : ~r2_hidden(X5,sK1) [avatar definition]
156. ~r2_hidden(X5,sK1) <- (10) [avatar component clause 155] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
158. 11 <=> ! [X4] : (~r2_hidden(X4,sK1) | r2_hidden(X4,sK0)) [avatar definition]
159. ~r2_hidden(X4,sK1) | r2_hidden(X4,sK0) <- (11) [avatar component clause 158] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
173. ~r2_hidden(X6,sK1) | ~r2_hidden(X7,sK1) | r2_hidden(X7,sK0) <- (2) [resolution 139,59] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
177. 11 | 10 | ~2 [avatar split clause 173,74,155,158]
187. r2_hidden(sK2(X0,sK0),sK1) | sK0 = X0 | r2_hidden(sK2(X0,sK0),X0) <- (9) [resolution 145,55] {a:3,w:13,wCS:1300,nSel:1,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
223. 16 <=> r2_hidden(sK2(sK1,sK0),sK1) [avatar definition]
224. r2_hidden(sK2(sK1,sK0),sK1) <- (16) [avatar component clause 223] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
260. r2_hidden(sK2(sK1,sK0),sK1) | sK0 = sK1 <- (9) [factoring 187] {a:4,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
261. r2_hidden(sK2(sK1,sK0),sK1) <- (~1, 9) [subsumption resolution 260,71] {a:4,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
262. 16 | 1 | ~9 [avatar split clause 261,144,69,223]
263. r2_hidden(sK2(sK1,sK0),sK0) <- (11, 16) [resolution 224,159] {a:5,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
267. 19 <=> r2_hidden(sK2(sK1,sK0),sK0) [avatar definition]
269. r2_hidden(sK2(sK1,sK0),sK0) <- (19) [avatar component clause 267] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
270. 19 | ~11 | ~16 [avatar split clause 263,223,158,267]
277. sK0 = sK1 | ~r2_hidden(sK2(sK1,sK0),sK1) <- (19) [resolution 269,56] {a:6,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
282. ~r2_hidden(sK2(sK1,sK0),sK1) <- (~1, 19) [subsumption resolution 277,71] {a:6,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
283. $false <- (~1, 16, 19) [subsumption resolution 282,224] {a:6,w:0,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
284. 1 | ~16 | ~19 [avatar contradiction clause 283]
290. $false <- (~1, 8, 10) [unit resulting resolution 71,142,156,55] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
292. $false <- (~1, 9, 10) [unit resulting resolution 71,156,156,187] {a:4,w:0,goal:1,sine:2,thAx:0,allAx:16,thDist:-16}
301. 1 | ~8 | ~10 [avatar contradiction clause 290]
303. 1 | ~9 | ~10 [avatar contradiction clause 292]
304. ~r2_hidden(X4,sK1) <- (8, 11) [subsumption resolution 159,142] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
305. 10 | ~8 | ~11 [avatar split clause 304,158,141,155]
306. $false [avatar sat refutation 72,77,147,177,262,270,284,301,303,305] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:307,thDist:-307}
% SZS output end Proof for zfmisc_1__t92_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 177
% Active clauses: 49
% Passive clauses: 98
% Final active clauses: 34
% Final passive clauses: 34
% Main loop iterations started: 61
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 7
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Equational tautologies: 2
% Forward subsumptions: 32
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 54
% Unit resulting resolution: 6
% Factoring: 2
% Forward superposition: 29
% Backward superposition: 22
% Self superposition: 6
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 26
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 37
% SAT solver unit clauses: 7
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9594
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% splitting model update: 0.001 s ( own 0.0000 s ) 
% backward subsumption index maintenance: 0.001 s
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

