i: [1,52,0,3,1,1,0,0] t10_taxonom2
a: [3,79,0,3,1,52]
s: 79
i: [1,53,0,2,1,0,0,255] fc1_xboole_0
a: [3,84,0,2,1,53]
s: 84
i: [1,71,0,2,1,0,0,2] rc2_xboole_0
a: [3,89,0,2,1,71]
s: 89
i: [1,72,0,2,1,0,0,2] rc1_xboole_0
a: [3,94,0,2,1,72]
s: 94
i: [1,54,0,3,1,0,0,1] fc2_xboole_0
s: 54
i: [1,68,0,4,1,0,0,255] existence_m1_subset_1
s: 68
i: [1,74,0,4,1,0,0,1] d1_tarski
s: 74
i: [1,55,0,5,2,0,0,255] t6_boole
s: 55
i: [1,67,0,5,2,0,0,2] t7_boole
s: 67
i: [1,56,0,6,2,0,0,1] d5_taxonom2
s: 56
i: [1,57,0,6,2,0,0,1] d5_taxonom2
s: 57
d: [2,99,1,4,2,0,40,56,67]
s: 99
i: [1,60,0,6,2,0,0,255] t1_subset
s: 60
d: [2,101,2,2,1,1,40,99,84]
a: [3,106,2,2,1,101]
s: 106
i: [1,69,0,6,2,0,0,2] symmetry_r1_xboole_0
s: 69
i: [1,70,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 70
d: [2,102,2,2,1,1,40,99,94]
a: [3,111,2,2,1,102]
s: 111
i: [1,58,0,7,2,0,0,1] d5_taxonom2
s: 58
d: [2,100,1,4,2,0,40,57,67]
s: 100
i: [1,59,0,7,2,0,0,1] d5_taxonom2
s: 59
i: [1,66,0,7,3,0,0,2] t8_boole
s: 66
d: [2,113,1,4,1,0,40,60,74]
s: 113
i: [1,75,0,7,2,0,0,1] d1_tarski
s: 75
d: [2,117,1,4,1,0,40,70,74]
s: 117
i: [1,61,0,8,3,0,0,255] t2_subset
s: 61
i: [1,64,0,14,3,0,0,1] d1_tarski
s: 64
d: [2,114,1,6,2,0,40,60,56]
s: 114
i: [1,65,0,14,3,0,0,1] d1_tarski
s: 65
d: [2,115,1,6,2,0,40,60,57]
s: 115
d: [2,118,1,6,2,0,40,70,56]
s: 118
d: [2,119,1,6,2,0,40,70,57]
s: 119
d: [2,123,1,6,2,0,40,61,68]
s: 123
d: [2,96,1,7,3,0,44,55,55]
s: 96
d: [2,134,2,8,2,0,40,123,75]
d: [2,135,2,5,1,0,21,134,54]
s: 135
d: [2,121,1,8,2,0,40,75,56]
s: 121
d: [2,122,1,8,2,0,40,75,57]
s: 122
d: [2,149,2,11,3,0,44,58,122]
d: [2,152,2,8,2,0,18,149]
d: [2,156,2,3,1,0,21,152,121]
s: 156
r: 121
r: 122
d: [2,158,3,0,0,1,40,156,79]
e: 158
d: [2,157,3,0,0,1,52,79,156]
e: 157
f: 157,52
% Refutation found. Thanks to Tanya!
% SZS status Theorem for taxonom2__t10_taxonom2
% SZS output start Proof for taxonom2__t10_taxonom2
1. ! [X0] : v4_taxonom2(k1_tarski(X0)) [input t10_taxonom2]
2. ~! [X0] : v4_taxonom2(k1_tarski(X0)) [negated conjecture 1]
4. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
5. ! [X0] : (v4_taxonom2(X0) <=> ! [X1] : ! [X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (r1_xboole_0(X1,X2) | X1 = X2))) [input d5_taxonom2]
20. ! [X0] : (v4_taxonom2(X0) <=> ! [X1,X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (r1_xboole_0(X1,X2) | X1 = X2))) [flattening 5]
23. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 4]
26. ! [X0] : (! [X1,X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (r1_xboole_0(X1,X2) | X1 = X2)) => v4_taxonom2(X0)) [unused predicate definition removal 20]
27. ? [X0] : ~v4_taxonom2(k1_tarski(X0)) [ennf transformation 2]
29. ! [X0] : (v4_taxonom2(X0) | ? [X1,X2] : ((~r1_xboole_0(X1,X2) & X1 != X2) & (r2_hidden(X2,X0) & r2_hidden(X1,X0)))) [ennf transformation 26]
30. ! [X0] : (v4_taxonom2(X0) | ? [X1,X2] : (~r1_xboole_0(X1,X2) & X1 != X2 & r2_hidden(X2,X0) & r2_hidden(X1,X0))) [flattening 29]
38. ? [X0] : ~v4_taxonom2(k1_tarski(X0)) => ~v4_taxonom2(k1_tarski(sK0)) [choice axiom]
39. ~v4_taxonom2(k1_tarski(sK0)) [skolemisation 27,38]
40. ! [X0] : (? [X1,X2] : (~r1_xboole_0(X1,X2) & X1 != X2 & r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (~r1_xboole_0(sK1(X0),sK2(X0)) & sK1(X0) != sK2(X0) & r2_hidden(sK2(X0),X0) & r2_hidden(sK1(X0),X0))) [choice axiom]
41. ! [X0] : (v4_taxonom2(X0) | (~r1_xboole_0(sK1(X0),sK2(X0)) & sK1(X0) != sK2(X0) & r2_hidden(sK2(X0),X0) & r2_hidden(sK1(X0),X0))) [skolemisation 30,40]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 23]
43. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 42]
44. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
45. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 43,44]
52. ~v4_taxonom2(k1_tarski(sK0)) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
56. r2_hidden(sK1(X0),X0) | v4_taxonom2(X0) [cnf transformation 41] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
57. r2_hidden(sK2(X0),X0) | v4_taxonom2(X0) [cnf transformation 41] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
58. sK1(X0) != sK2(X0) | v4_taxonom2(X0) [cnf transformation 41] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
62. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 45] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
75. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 62] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
77. 1 <=> v4_taxonom2(k1_tarski(sK0)) [avatar definition]
79. ~v4_taxonom2(k1_tarski(sK0)) <- (~1) [avatar component clause 77] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. ~1 [avatar split clause 52,77]
121. sK1(k1_tarski(X1)) = X1 | v4_taxonom2(k1_tarski(X1)) [resolution 75,56] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
122. sK2(k1_tarski(X2)) = X2 | v4_taxonom2(k1_tarski(X2)) [resolution 75,57] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
149. sK1(k1_tarski(X3)) != X3 | v4_taxonom2(k1_tarski(X3)) | v4_taxonom2(k1_tarski(X3)) [superposition 58,122] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:6,thDist:-6}
152. sK1(k1_tarski(X3)) != X3 | v4_taxonom2(k1_tarski(X3)) [duplicate literal removal 149] {a:2,w:8,wCS:3200,sine:1,thAx:0,allAx:6,thDist:-6}
156. v4_taxonom2(k1_tarski(X3)) [subsumption resolution 152,121] {a:2,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
157. $false <- (~1) [unit resulting resolution 79,156] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
160. 1 [avatar contradiction clause 157]
161. $false [avatar sat refutation 80,160] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:40,thDist:-40}
% SZS output end Proof for taxonom2__t10_taxonom2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% >>> Preprocessing
% Introduced skolems: 7
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 80
% Active clauses: 35
% Passive clauses: 44
% Final active clauses: 33
% Final passive clauses: 7
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 1
% Forward subsumptions: 16
% Backward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 27
% Unit resulting resolution: 1
% Forward superposition: 1
% Backward superposition: 12
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 6
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

