i: [1,42,0,2,1,1,0,0] t1_xregular
a: [3,62,0,2,1,42]
s: 62
i: [1,43,0,6,2,1,0,0] t1_xregular
s: 43
i: [1,44,0,2,1,0,0,255] fc1_xboole_0
a: [3,67,0,2,1,44]
s: 67
i: [1,57,0,2,1,0,0,1] rc2_xboole_0
a: [3,72,0,2,1,57]
s: 72
i: [1,58,0,2,1,0,0,1] rc1_xboole_0
a: [3,77,0,2,1,58]
s: 77
i: [1,45,0,5,2,0,0,255] t6_boole
s: 45
i: [1,46,0,5,2,0,0,1] d1_xboole_0
s: 46
i: [1,52,0,5,2,0,0,1] t7_boole
s: 52
i: [1,47,0,6,2,0,0,1] d1_xboole_0
s: 47
i: [1,55,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 55
i: [1,56,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 56
d: [2,82,1,6,2,0,40,56,47]
s: 82
i: [1,51,0,7,3,0,0,255] t8_boole
s: 51
i: [1,53,0,7,2,0,0,1] t2_tarski
s: 53
i: [1,48,0,8,2,0,0,1] t3_xboole_0
s: 48
i: [1,49,0,8,2,0,0,1] t3_xboole_0
s: 49
d: [2,86,1,5,2,0,40,48,46]
s: 86
i: [1,50,0,9,3,0,0,1] t3_xboole_0
s: 50
d: [2,89,2,5,2,0,40,86,43]
s: 89
i: [1,54,0,10,3,0,0,1] t2_tarski
s: 54
d: [2,88,1,5,2,0,40,49,46]
s: 88
d: [2,92,3,5,2,0,40,89,47]
d: [2,96,3,3,1,1,21,92,62]
a: [3,100,3,3,1,96]
s: 100
d: [2,79,1,7,3,0,44,45,45]
s: 79
d: [2,93,3,6,2,0,40,89,53]
a: [3,108,3,3,1,93]
s: 108
d: [2,83,1,7,2,0,40,53,56]
s: 83
d: [2,85,1,8,2,0,40,48,56]
s: 85
d: [2,94,3,7,2,0,40,89,48]
s: 94
d: [2,87,1,8,2,0,40,49,56]
s: 87
d: [2,95,3,7,2,0,40,89,49]
s: 95
d: [2,110,1,11,3,0,40,54,47]
s: 110
d: [2,111,1,12,3,0,40,54,53]
s: 111
d: [2,90,2,5,2,0,40,86,55]
s: 90
d: [2,112,1,13,3,0,40,54,48]
s: 112
d: [2,117,2,11,3,0,40,112,49]
d: [2,118,2,7,2,0,18,117]
s: 118
d: [2,113,1,13,3,0,40,54,49]
s: 113
d: [2,124,2,11,3,0,40,113,48]
d: [2,125,2,7,2,0,18,124]
s: 125
a: [3,104,3,3,1,93]
s: 104
r: 43
r: 89
d: [2,128,4,0,0,2,52,62,104,47]
e: 128
d: [2,129,4,2,1,1,40,104,47]
d: [2,134,4,0,0,2,21,129,62]
e: 134
d: [2,119,3,7,2,0,40,118,43]
d: [2,122,3,3,1,0,21,119,53]
f: 134,122,42
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xregular__t1_xregular
% SZS output start Proof for xregular__t1_xregular
1. ! [X0] : (~v1_xboole_0(X0) => ? [X1] : (r1_xboole_0(X1,X0) & r2_hidden(X1,X0))) [input t1_xregular]
2. ~! [X0] : (~v1_xboole_0(X0) => ? [X1] : (r1_xboole_0(X1,X0) & r2_hidden(X1,X0))) [negated conjecture 1]
4. ! [X0] : (v1_xboole_0(X0) <=> ! [X1] : ~r2_hidden(X1,X0)) [input d1_xboole_0]
10. ! [X0] : ! [X1] : ~(! [X2] : ~(! [X3] : ~(r2_hidden(X3,X2) & r2_hidden(X3,X1)) & r2_hidden(X2,X1)) & r2_hidden(X0,X1)) [input t2_tarski]
11. ! [X0] : ! [X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X2] : ~(r2_hidden(X2,X1) & r2_hidden(X2,X0)) & ~r1_xboole_0(X0,X1))) [input t3_xboole_0]
15. ! [X0] : ! [X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X3] : ~(r2_hidden(X3,X1) & r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) [rectify 11]
16. ! [X0,X1] : (~(r1_xboole_0(X0,X1) & ? [X2] : (r2_hidden(X2,X1) & r2_hidden(X2,X0))) & ~(! [X3] : ~(r2_hidden(X3,X1) & r2_hidden(X3,X0)) & ~r1_xboole_0(X0,X1))) [flattening 15]
19. ! [X0,X1] : ~(! [X2] : ~(! [X3] : ~(r2_hidden(X3,X2) & r2_hidden(X3,X1)) & r2_hidden(X2,X1)) & r2_hidden(X0,X1)) [flattening 10]
20. ? [X0] : (! [X1] : (~r1_xboole_0(X1,X0) | ~r2_hidden(X1,X0)) & ~v1_xboole_0(X0)) [ennf transformation 2]
22. ! [X0,X1] : ((~r1_xboole_0(X0,X1) | ! [X2] : (~r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) & (? [X3] : (r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r1_xboole_0(X0,X1))) [ennf transformation 16]
25. ! [X0,X1] : (? [X2] : (! [X3] : (~r2_hidden(X3,X2) | ~r2_hidden(X3,X1)) & r2_hidden(X2,X1)) | ~r2_hidden(X0,X1)) [ennf transformation 19]
28. ? [X0] : (! [X1] : (~r1_xboole_0(X1,X0) | ~r2_hidden(X1,X0)) & ~v1_xboole_0(X0)) => (! [X1] : (~r1_xboole_0(X1,sK0) | ~r2_hidden(X1,sK0)) & ~v1_xboole_0(sK0)) [choice axiom]
29. ! [X1] : (~r1_xboole_0(X1,sK0) | ~r2_hidden(X1,sK0)) & ~v1_xboole_0(sK0) [skolemisation 20,28]
30. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X1] : ~r2_hidden(X1,X0) | ~v1_xboole_0(X0))) [nnf transformation 4]
31. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [rectify 30]
32. ! [X0] : (? [X1] : r2_hidden(X1,X0) => r2_hidden(sK1(X0),X0)) [choice axiom]
33. ! [X0] : ((v1_xboole_0(X0) | r2_hidden(sK1(X0),X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [skolemisation 31,32]
34. ! [X1,X0] : (? [X3] : (r2_hidden(X3,X1) & r2_hidden(X3,X0)) => (r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) [choice axiom]
35. ! [X0,X1] : ((~r1_xboole_0(X0,X1) | ! [X2] : (~r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) & ((r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0)) | r1_xboole_0(X0,X1))) [skolemisation 22,34]
36. ! [X1] : (? [X2] : (! [X3] : (~r2_hidden(X3,X2) | ~r2_hidden(X3,X1)) & r2_hidden(X2,X1)) => (! [X3] : (~r2_hidden(X3,sK3(X1)) | ~r2_hidden(X3,X1)) & r2_hidden(sK3(X1),X1))) [choice axiom]
37. ! [X0,X1] : ((! [X3] : (~r2_hidden(X3,sK3(X1)) | ~r2_hidden(X3,X1)) & r2_hidden(sK3(X1),X1)) | ~r2_hidden(X0,X1)) [skolemisation 25,36]
42. ~v1_xboole_0(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. ~r1_xboole_0(X1,sK0) | ~r2_hidden(X1,sK0) [cnf transformation 29] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. r2_hidden(sK1(X0),X0) | v1_xboole_0(X0) [cnf transformation 33] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
48. r2_hidden(sK2(X0,X1),X0) | r1_xboole_0(X0,X1) [cnf transformation 35] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
49. r2_hidden(sK2(X0,X1),X1) | r1_xboole_0(X0,X1) [cnf transformation 35] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
53. r2_hidden(sK3(X1),X1) | ~r2_hidden(X0,X1) [cnf transformation 37] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
54. ~r2_hidden(X3,sK3(X1)) | ~r2_hidden(X3,X1) | ~r2_hidden(X0,X1) [cnf transformation 37] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
60. 1 <=> v1_xboole_0(sK0) [avatar definition]
62. ~v1_xboole_0(sK0) <- (~1) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. ~1 [avatar split clause 42,60]
103. 6 <=> ! [X0] : ~r2_hidden(X0,sK0) [avatar definition]
104. ~r2_hidden(X0,sK0) <- (6) [avatar component clause 103] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
112. ~r2_hidden(sK2(sK3(X5),X6),X5) | ~r2_hidden(X7,X5) | r1_xboole_0(sK3(X5),X6) [resolution 54,48] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
117. ~r2_hidden(X0,X1) | r1_xboole_0(sK3(X1),X1) | r1_xboole_0(sK3(X1),X1) [resolution 112,49] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:6,thDist:-6}
118. r1_xboole_0(sK3(X1),X1) | ~r2_hidden(X0,X1) [duplicate literal removal 117] {a:2,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
119. ~r2_hidden(X0,sK0) | ~r2_hidden(sK3(sK0),sK0) [resolution 118,43] {a:3,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
122. ~r2_hidden(X0,sK0) [subsumption resolution 119,53] {a:3,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
123. 6 [avatar split clause 122,103]
129. v1_xboole_0(sK0) <- (6) [resolution 104,47] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
134. $false <- (~1, 6) [subsumption resolution 129,62] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
135. 1 | ~6 [avatar contradiction clause 134]
136. $false [avatar sat refutation 63,123,135] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:62,thDist:-62}
% SZS output end Proof for xregular__t1_xregular
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 65
% Active clauses: 34
% Passive clauses: 38
% Final active clauses: 32
% Final passive clauses: 1
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 12
% Backward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 33
% Unit resulting resolution: 1
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 5
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

