i: [1,125,0,2,1,1,0,0] t35_ordinal3
a: [3,196,0,2,1,125]
s: 196
i: [1,124,0,2,1,1,0,0] t35_ordinal3
a: [3,201,0,2,1,124]
s: 201
i: [1,123,0,2,1,1,0,0] t35_ordinal3
a: [3,206,0,2,1,123]
s: 206
i: [1,128,0,3,1,1,0,0] t35_ordinal3
a: [3,181,0,3,1,128]
s: 181
i: [1,127,0,3,1,1,0,0] t35_ordinal3
a: [3,186,0,3,1,127]
s: 186
i: [1,126,0,7,1,1,0,0] t35_ordinal3
a: [3,191,0,7,1,126]
s: 191
i: [1,129,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,211,0,2,1,129]
s: 211
i: [1,130,0,2,1,0,0,1] fc1_xboole_0
a: [3,216,0,2,1,130]
s: 216
i: [1,168,0,2,1,0,0,2] rc2_xboole_0
a: [3,233,0,2,1,168]
s: 233
i: [1,169,0,2,1,0,0,255] rc4_ordinal1
a: [3,238,0,2,1,169]
s: 238
i: [1,170,0,2,1,0,0,2] rc1_xboole_0
a: [3,243,0,2,1,170]
s: 243
i: [1,171,0,2,1,0,0,1] rc1_ordinal1
a: [3,248,0,2,1,171]
s: 248
i: [1,175,0,2,1,0,0,255] rc2_ordinal1
a: [3,253,0,2,1,175]
s: 253
i: [1,174,0,2,1,0,0,255] rc2_ordinal1
a: [3,258,0,2,1,174]
s: 258
i: [1,173,0,2,1,0,0,255] rc2_ordinal1
a: [3,263,0,2,1,173]
s: 263
i: [1,172,0,2,1,0,0,255] rc2_ordinal1
a: [3,268,0,2,1,172]
s: 268
i: [1,177,0,2,1,0,0,255] rc5_ordinal1
a: [3,273,0,2,1,177]
s: 273
i: [1,176,0,2,1,0,0,255] rc5_ordinal1
a: [3,278,0,2,1,176]
s: 278
i: [1,159,0,3,1,0,0,255] irreflexivity_r2_xboole_0
s: 159
i: [1,160,0,3,1,0,0,2] reflexivity_r1_tarski
s: 160
i: [1,131,0,3,1,0,0,255] d2_xboole_0
a: [3,221,0,3,1,131]
s: 221
r: 211
i: [1,134,0,4,2,0,0,255] cc6_ordinal1
s: 134
i: [1,135,0,4,2,0,0,255] cc9_ordinal1
s: 135
d: [2,281,1,2,1,1,40,134,238]
a: [3,286,1,2,1,281]
s: 286
i: [1,136,0,4,2,0,0,255] cc7_ordinal1
s: 136
i: [1,137,0,4,2,0,0,2] cc3_ordinal1
s: 137
d: [2,282,1,2,1,1,40,134,273]
a: [3,291,1,2,1,282]
s: 291
i: [1,139,0,4,2,0,0,255] cc1_ordinal1
s: 139
d: [2,329,1,2,1,1,40,139,206]
a: [3,339,1,2,1,329]
s: 339
i: [1,140,0,4,2,0,0,255] cc1_ordinal1
s: 140
i: [1,158,0,4,1,0,0,3] existence_m1_subset_1
s: 158
d: [2,330,1,2,1,1,40,139,201]
a: [3,344,1,2,1,330]
s: 344
i: [1,138,0,5,2,0,0,1] t6_boole
s: 138
d: [2,331,1,2,1,1,40,139,196]
a: [3,349,1,2,1,331]
s: 349
i: [1,155,0,5,2,0,0,2] t7_boole
s: 155
i: [1,163,0,7,3,0,0,1] reflexivity_r1_ordinal1
a: [3,228,0,5,2,163]
s: 228
d: [2,366,1,2,1,1,40,140,206]
a: [3,376,1,2,1,366]
s: 376
i: [1,145,0,6,3,0,0,255] cc2_ordinal1
s: 145
d: [2,367,1,2,1,1,40,140,201]
a: [3,381,1,2,1,367]
s: 381
i: [1,147,0,6,2,0,0,2] d2_ordinal1
s: 147
i: [1,148,0,6,2,0,0,2] d2_ordinal1
s: 148
d: [2,368,1,2,1,1,40,140,196]
a: [3,386,1,2,1,368]
s: 386
i: [1,150,0,6,2,0,0,2] t1_subset
s: 150
d: [2,293,1,2,1,1,40,135,243]
a: [3,298,1,2,1,293]
s: 298
i: [1,161,0,6,2,0,0,255] antisymmetry_r2_xboole_0
s: 161
i: [1,162,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 162
d: [2,294,1,2,1,1,40,135,216]
a: [3,303,1,2,1,294]
s: 303
i: [1,143,0,7,3,0,0,3] cc5_ordinal1
s: 143
d: [2,305,1,2,1,1,40,136,243]
a: [3,310,1,2,1,305]
s: 310
i: [1,144,0,7,3,0,0,3] cc2_ordinal2
s: 144
i: [1,152,0,7,2,0,0,255] t3_subset
s: 152
d: [2,306,1,2,1,1,40,136,216]
a: [3,315,1,2,1,306]
s: 315
i: [1,153,0,7,2,0,0,255] t3_subset
s: 153
d: [2,317,1,2,1,1,40,137,243]
a: [3,322,1,2,1,317]
s: 322
i: [1,154,0,7,3,0,0,1] t8_boole
s: 154
i: [1,132,0,8,3,0,0,255] cc10_ordinal1
s: 132
d: [2,318,1,2,1,1,40,137,216]
a: [3,327,1,2,1,318]
s: 327
i: [1,146,0,8,3,0,0,2] d2_ordinal1
s: 146
d: [2,333,1,2,1,1,40,139,248]
a: [3,359,1,2,1,333]
s: 359
i: [1,149,0,8,3,0,0,2] t2_subset
s: 149
i: [1,164,0,8,3,0,0,1] dt_k11_ordinal2
s: 164
d: [2,370,1,2,1,1,40,140,248]
a: [3,396,1,2,1,370]
s: 396
i: [1,151,0,9,3,0,0,255] d8_xboole_0
s: 151
d: [2,332,2,2,1,1,40,139,286]
a: [3,354,2,2,1,332]
s: 354
i: [1,157,0,9,3,0,0,255] t5_subset
s: 157
i: [1,133,0,10,4,0,0,255] t11_ordinal1
s: 133
d: [2,335,2,2,1,1,40,139,291]
a: [3,364,2,2,1,335]
s: 364
i: [1,156,0,10,3,0,0,255] t4_subset
s: 156
d: [2,369,2,2,1,1,40,140,286]
a: [3,391,2,2,1,369]
s: 391
i: [1,165,0,10,4,0,0,1] connectedness_r1_ordinal1
s: 165
i: [1,166,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 166
d: [2,461,1,7,3,1,40,165,181]
d: [2,464,1,5,2,2,21,461,206]
d: [2,465,1,3,1,3,21,464,196]
a: [3,469,1,3,1,465]
s: 469
i: [1,167,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 167
d: [2,489,2,7,3,1,40,469,166]
d: [2,490,2,5,2,2,21,489,206]
d: [2,491,2,3,1,3,21,490,196]
a: [3,495,2,3,1,491]
s: 495
i: [1,142,0,16,5,0,0,1] t34_ordinal3
s: 142
i: [1,141,0,19,6,0,0,1] t33_ordinal3
s: 141
d: [2,500,3,6,2,1,40,495,151]
a: [3,504,3,3,1,500]
s: 504
d: [2,471,1,15,3,1,40,166,191]
a: [3,479,1,4,1,471]
s: 479
d: [2,527,2,0,0,3,52,206,201,479,164]
e: 527
d: [2,528,2,4,2,1,40,479,164]
d: [2,530,2,2,1,2,21,528,201]
d: [2,531,2,0,0,3,21,530,206]
e: 531
r: 479
a: [3,507,3,3,1,500]
s: 507
a: [3,478,1,4,1,471]
s: 478
a: [3,483,1,4,1,471]
s: 483
d: [2,545,2,0,0,3,52,196,201,483,164]
e: 545
d: [2,546,2,4,2,1,40,483,164]
d: [2,548,2,2,1,2,21,546,201]
d: [2,549,2,0,0,3,21,548,196]
e: 549
r: 483
d: [2,512,4,7,3,1,40,504,133]
d: [2,514,4,5,2,2,21,512,206]
d: [2,515,4,3,1,3,21,514,349]
a: [3,519,4,3,1,515]
s: 519
a: [3,482,1,4,1,471]
s: 482
d: [2,554,5,2,1,1,40,519,155]
a: [3,568,5,2,1,554]
s: 568
a: [3,487,1,7,1,471]
s: 487
d: [2,408,1,4,2,0,40,147,155]
s: 408
d: [2,513,4,3,1,1,40,504,161]
a: [3,524,4,3,1,513]
s: 524
d: [2,413,1,5,2,0,40,143,158]
s: 413
d: [2,552,5,3,1,1,40,519,162]
a: [3,558,5,3,1,552]
s: 558
d: [2,416,1,5,1,0,40,152,158]
s: 416
d: [2,410,1,6,2,0,40,150,147]
s: 410
d: [2,553,5,3,1,1,40,519,150]
a: [3,563,5,3,1,553]
s: 563
d: [2,411,1,6,2,0,40,162,147]
s: 411
d: [2,533,2,4,1,1,40,478,140]
a: [3,538,2,4,1,533]
s: 538
d: [2,434,1,6,2,0,40,132,158]
s: 434
d: [2,450,1,6,2,0,40,149,158]
s: 450
d: [2,534,2,4,1,1,40,478,139]
a: [3,543,2,4,1,534]
s: 543
d: [2,403,1,7,3,0,44,138,138]
s: 403
d: [2,570,2,4,1,1,40,482,140]
a: [3,575,2,4,1,570]
s: 575
d: [2,435,1,7,3,0,40,132,153]
s: 435
d: [2,456,1,7,2,0,40,157,158]
s: 456
d: [2,617,3,4,2,1,40,435,495]
a: [3,635,3,2,1,617]
s: 635
d: [2,421,1,8,3,0,40,153,143]
s: 421
d: [2,571,2,4,1,1,40,482,139]
a: [3,580,2,4,1,571]
s: 580
d: [2,452,1,8,3,0,40,164,140]
s: 452
d: [2,453,1,8,3,0,40,164,139]
s: 453
d: [2,616,2,8,2,1,40,435,487]
a: [3,626,2,4,1,616]
s: 626
d: [2,457,1,8,3,0,40,157,153]
s: 457
d: [2,645,2,9,2,1,40,457,487]
a: [3,651,2,4,1,645]
s: 651
d: [2,458,1,8,2,0,40,156,158]
s: 458
d: [2,459,1,9,3,0,40,156,153]
s: 459
d: [2,660,3,6,2,1,40,459,495]
s: 660
d: [2,451,1,10,3,0,40,149,153]
s: 451
d: [2,663,4,6,2,1,40,660,450]
a: [3,667,4,2,1,663]
s: 667
r: 667
d: [2,473,1,14,6,0,40,166,165]
d: [2,474,1,10,4,0,18,473]
s: 474
d: [2,583,2,14,2,1,40,487,151]
a: [3,587,2,7,1,583]
s: 587
r: 587
a: [3,666,4,2,1,663]
s: 666
a: [3,586,2,7,1,583]
s: 586
a: [3,671,4,4,1,663]
s: 671
a: [3,591,2,7,1,583]
s: 591
r: 191
r: 482
r: 487
r: 575
r: 580
r: 586
r: 626
r: 651
d: [2,732,4,0,0,6,52,206,201,196,507,186,591,141]
e: 732
d: [2,733,4,0,0,6,52,196,201,206,507,186,591,141]
e: 733
r: 591
s: 191
d: [2,711,5,5,2,1,40,671,143]
d: [2,718,5,3,1,2,21,711,206]
a: [3,722,5,3,1,718]
s: 722
s: 482
d: [2,756,6,3,1,1,40,722,140]
a: [3,761,6,3,1,756]
s: 761
s: 487
s: 575
d: [2,757,6,3,1,1,40,722,139]
a: [3,766,6,3,1,757]
s: 766
s: 580
s: 626
s: 651
s: 587
d: [2,710,5,6,2,1,40,671,149]
d: [2,712,5,4,1,2,21,710,568]
a: [3,716,5,4,1,712]
s: 716
r: 666
r: 671
r: 716
a: [3,590,2,7,1,583]
s: 590
a: [3,631,3,2,1,617]
s: 631
d: [2,372,2,2,1,1,40,140,291]
a: [3,401,2,2,1,372]
s: 401
d: [2,422,2,2,1,1,40,322,140]
a: [3,427,2,2,1,422]
s: 427
s: 667
d: [2,423,2,2,1,1,40,322,139]
a: [3,432,2,2,1,423]
s: 432
d: [2,679,5,2,1,1,40,667,136]
a: [3,684,5,2,1,679]
s: 684
d: [2,436,2,2,1,1,40,327,140]
a: [3,442,2,2,1,436]
s: 442
d: [2,437,2,2,1,1,40,327,139]
a: [3,447,2,2,1,437]
s: 447
a: [3,670,4,4,1,663]
s: 670
d: [2,659,2,10,2,1,40,459,487]
s: 659
d: [2,816,3,12,2,1,40,659,450]
a: [3,820,3,4,1,816]
s: 820
r: 820
d: [2,773,2,10,2,1,40,487,459]
s: 773
d: [2,597,2,5,2,0,40,413,140]
s: 597
a: [3,819,3,4,1,816]
s: 819
d: [2,598,2,5,2,0,40,413,139]
s: 598
a: [3,715,5,4,1,712]
s: 715
d: [2,639,2,5,2,0,40,421,160]
s: 639
d: [2,611,2,6,2,0,40,450,162]
s: 611
d: [2,696,3,15,3,1,40,587,133]
d: [2,698,3,11,2,2,21,696,482]
d: [2,699,3,7,1,3,21,698,543]
a: [3,703,3,7,1,699]
s: 703
d: [2,837,6,0,0,5,52,206,201,196,558,703,142]
e: 837
d: [2,838,4,9,4,1,40,703,142]
d: [2,844,4,6,3,2,21,838,558]
d: [2,845,4,4,2,3,21,844,196]
d: [2,846,4,2,1,4,21,845,201]
d: [2,847,4,0,0,5,21,846,206]
e: 847
d: [2,753,1,15,3,1,40,191,166]
d: [2,754,1,11,2,2,21,753,478]
d: [2,778,2,14,2,1,40,487,151]
d: [2,783,3,15,3,1,40,587,133]
d: [2,785,3,11,2,2,21,783,482]
d: [2,786,3,7,1,3,21,785,543]
f: 847,786,778,754,733,552,549,534,531,515,500,491,465,331,123,124,125,126,127,128
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t35_ordinal3
% SZS output start Proof for ordinal3__t35_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r1_ordinal1(k11_ordinal2(X0,X1),k11_ordinal2(X2,X1)) => (r1_ordinal1(X0,X2) | k1_xboole_0 = X1))))) [input t35_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r1_ordinal1(k11_ordinal2(X0,X1),k11_ordinal2(X2,X1)) => (r1_ordinal1(X0,X2) | k1_xboole_0 = X1))))) [negated conjecture 1]
3. ! [X0,X1] : (r2_hidden(X0,X1) => ~r2_hidden(X1,X0)) [input antisymmetry_r2_hidden]
6. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
15. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1))) [input connectedness_r1_ordinal1]
18. ! [X0] : ! [X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [input d8_xboole_0]
19. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k11_ordinal2(X0,X1))) [input dt_k11_ordinal2]
33. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
36. ! [X0] : (v1_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_xboole_0(X0,X1) => r2_hidden(X0,X1)))) [input t11_ordinal1]
39. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (k11_ordinal2(X0,X1) = k11_ordinal2(X2,X1) => (X0 = X2 | k1_xboole_0 = X1))))) [input t33_ordinal3]
40. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r2_hidden(k11_ordinal2(X0,X1),k11_ordinal2(X2,X1)) => r2_hidden(X0,X2))))) [input t34_ordinal3]
49. ! [X0,X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [flattening 18]
57. ! [X0,X1] : ((X0 != X1 & r1_tarski(X0,X1)) => r2_xboole_0(X0,X1)) [unused predicate definition removal 49]
59. ? [X0] : (? [X1] : (? [X2] : (((~r1_ordinal1(X0,X2) & k1_xboole_0 != X1) & r1_ordinal1(k11_ordinal2(X0,X1),k11_ordinal2(X2,X1))) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
60. ? [X0] : (? [X1] : (? [X2] : (~r1_ordinal1(X0,X2) & k1_xboole_0 != X1 & r1_ordinal1(k11_ordinal2(X0,X1),k11_ordinal2(X2,X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 59]
62. ! [X0] : (! [X1] : ((r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1)) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [ennf transformation 36]
63. ! [X0] : (! [X1] : (r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [flattening 62]
69. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 6]
70. ! [X0] : (! [X1] : (! [X2] : (((X0 = X2 | k1_xboole_0 = X1) | k11_ordinal2(X0,X1) != k11_ordinal2(X2,X1)) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 39]
71. ! [X0] : (! [X1] : (! [X2] : (X0 = X2 | k1_xboole_0 = X1 | k11_ordinal2(X0,X1) != k11_ordinal2(X2,X1) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 70]
72. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X0,X2) | ~r2_hidden(k11_ordinal2(X0,X1),k11_ordinal2(X2,X1))) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 40]
73. ! [X0] : (! [X1] : (! [X2] : (r2_hidden(X0,X2) | ~r2_hidden(k11_ordinal2(X0,X1),k11_ordinal2(X2,X1)) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 72]
82. ! [X0,X1] : (r2_xboole_0(X0,X1) | (X0 = X1 | ~r1_tarski(X0,X1))) [ennf transformation 57]
83. ! [X0,X1] : (r2_xboole_0(X0,X1) | X0 = X1 | ~r1_tarski(X0,X1)) [flattening 82]
90. ! [X0,X1] : (~r2_hidden(X1,X0) | ~r2_hidden(X0,X1)) [ennf transformation 3]
93. ! [X0,X1] : (v3_ordinal1(k11_ordinal2(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 19]
94. ! [X0,X1] : (v3_ordinal1(k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 93]
95. ! [X0,X1] : ((r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 15]
96. ! [X0,X1] : (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 95]
97. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 33]
98. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 97]
99. ? [X0] : (? [X1] : (? [X2] : (~r1_ordinal1(X0,X2) & k1_xboole_0 != X1 & r1_ordinal1(k11_ordinal2(X0,X1),k11_ordinal2(X2,X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (? [X2] : (~r1_ordinal1(sK0,X2) & k1_xboole_0 != X1 & r1_ordinal1(k11_ordinal2(sK0,X1),k11_ordinal2(X2,X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
100. ? [X1] : (? [X2] : (~r1_ordinal1(sK0,X2) & k1_xboole_0 != X1 & r1_ordinal1(k11_ordinal2(sK0,X1),k11_ordinal2(X2,X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) => (? [X2] : (~r1_ordinal1(sK0,X2) & k1_xboole_0 != sK1 & r1_ordinal1(k11_ordinal2(sK0,sK1),k11_ordinal2(X2,sK1)) & v3_ordinal1(X2)) & v3_ordinal1(sK1)) [choice axiom]
101. ? [X2] : (~r1_ordinal1(sK0,X2) & k1_xboole_0 != sK1 & r1_ordinal1(k11_ordinal2(sK0,sK1),k11_ordinal2(X2,sK1)) & v3_ordinal1(X2)) => (~r1_ordinal1(sK0,sK2) & k1_xboole_0 != sK1 & r1_ordinal1(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) & v3_ordinal1(sK2)) [choice axiom]
102. ((~r1_ordinal1(sK0,sK2) & k1_xboole_0 != sK1 & r1_ordinal1(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) & v3_ordinal1(sK2)) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 60,101,100,99]
110. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 98]
123. v3_ordinal1(sK0) [cnf transformation 102] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
124. v3_ordinal1(sK1) [cnf transformation 102] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
125. v3_ordinal1(sK2) [cnf transformation 102] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
126. r1_ordinal1(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) [cnf transformation 102] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
127. k1_xboole_0 != sK1 [cnf transformation 102] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
128. ~r1_ordinal1(sK0,sK2) [cnf transformation 102] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
133. ~r2_xboole_0(X0,X1) | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v1_ordinal1(X0) [cnf transformation 63] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
139. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 69] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
141. k11_ordinal2(X0,X1) != k11_ordinal2(X2,X1) | k1_xboole_0 = X1 | X0 = X2 | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 71] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
142. ~r2_hidden(k11_ordinal2(X0,X1),k11_ordinal2(X2,X1)) | r2_hidden(X0,X2) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 73] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
151. ~r1_tarski(X0,X1) | X0 = X1 | r2_xboole_0(X0,X1) [cnf transformation 83] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
162. ~r2_hidden(X1,X0) | ~r2_hidden(X0,X1) [cnf transformation 90] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
164. v3_ordinal1(k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 94] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
165. r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 96] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
166. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 110] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
179. 1 <=> r1_ordinal1(sK0,sK2) [avatar definition]
181. ~r1_ordinal1(sK0,sK2) <- (~1) [avatar component clause 179] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
182. ~1 [avatar split clause 128,179]
184. 2 <=> k1_xboole_0 = sK1 [avatar definition]
186. k1_xboole_0 != sK1 <- (~2) [avatar component clause 184] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
187. ~2 [avatar split clause 127,184]
189. 3 <=> r1_ordinal1(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) [avatar definition]
191. r1_ordinal1(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) <- (3) [avatar component clause 189] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
192. 3 [avatar split clause 126,189]
194. 4 <=> v3_ordinal1(sK2) [avatar definition]
196. v3_ordinal1(sK2) <- (4) [avatar component clause 194] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
197. 4 [avatar split clause 125,194]
199. 5 <=> v3_ordinal1(sK1) [avatar definition]
201. v3_ordinal1(sK1) <- (5) [avatar component clause 199] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
202. 5 [avatar split clause 124,199]
204. 6 <=> v3_ordinal1(sK0) [avatar definition]
206. v3_ordinal1(sK0) <- (6) [avatar component clause 204] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
207. 6 [avatar split clause 123,204]
331. v1_ordinal1(sK2) <- (4) [resolution 139,196] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
347. 32 <=> v1_ordinal1(sK2) [avatar definition]
349. v1_ordinal1(sK2) <- (32) [avatar component clause 347] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
350. 32 | ~4 [avatar split clause 331,194,347]
461. r1_ordinal1(sK2,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK2) <- (~1) [resolution 165,181] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
464. r1_ordinal1(sK2,sK0) | ~v3_ordinal1(sK2) <- (~1, 6) [subsumption resolution 461,206] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
465. r1_ordinal1(sK2,sK0) <- (~1, 4, 6) [subsumption resolution 464,196] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
467. 46 <=> r1_ordinal1(sK2,sK0) [avatar definition]
469. r1_ordinal1(sK2,sK0) <- (46) [avatar component clause 467] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
470. 46 | 1 | ~4 | ~6 [avatar split clause 465,204,194,179,467]
477. 47 <=> v3_ordinal1(k11_ordinal2(sK0,sK1)) [avatar definition]
478. v3_ordinal1(k11_ordinal2(sK0,sK1)) <- (47) [avatar component clause 477] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
479. ~v3_ordinal1(k11_ordinal2(sK0,sK1)) <- (~47) [avatar component clause 477] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
481. 48 <=> v3_ordinal1(k11_ordinal2(sK2,sK1)) [avatar definition]
482. v3_ordinal1(k11_ordinal2(sK2,sK1)) <- (48) [avatar component clause 481] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
483. ~v3_ordinal1(k11_ordinal2(sK2,sK1)) <- (~48) [avatar component clause 481] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
485. 49 <=> r1_tarski(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) [avatar definition]
487. r1_tarski(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) <- (49) [avatar component clause 485] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
489. r1_tarski(sK2,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK2) <- (46) [resolution 469,166] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
490. r1_tarski(sK2,sK0) | ~v3_ordinal1(sK2) <- (6, 46) [subsumption resolution 489,206] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
491. r1_tarski(sK2,sK0) <- (4, 6, 46) [subsumption resolution 490,196] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
493. 50 <=> r1_tarski(sK2,sK0) [avatar definition]
495. r1_tarski(sK2,sK0) <- (50) [avatar component clause 493] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
496. 50 | ~4 | ~6 | ~46 [avatar split clause 491,467,204,194,493]
500. sK0 = sK2 | r2_xboole_0(sK2,sK0) <- (50) [resolution 495,151] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
502. 51 <=> r2_xboole_0(sK2,sK0) [avatar definition]
504. r2_xboole_0(sK2,sK0) <- (51) [avatar component clause 502] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
506. 52 <=> sK0 = sK2 [avatar definition]
507. sK0 != sK2 <- (~52) [avatar component clause 506] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
509. 51 | 52 | ~50 [avatar split clause 500,493,506,502]
512. r2_hidden(sK2,sK0) | ~v3_ordinal1(sK0) | ~v1_ordinal1(sK2) <- (51) [resolution 504,133] {a:4,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
514. r2_hidden(sK2,sK0) | ~v1_ordinal1(sK2) <- (6, 51) [subsumption resolution 512,206] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
515. r2_hidden(sK2,sK0) <- (6, 32, 51) [subsumption resolution 514,349] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
517. 53 <=> r2_hidden(sK2,sK0) [avatar definition]
519. r2_hidden(sK2,sK0) <- (53) [avatar component clause 517] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
520. 53 | ~6 | ~32 | ~51 [avatar split clause 515,502,347,204,517]
526. sK0 != sK2 | r1_ordinal1(sK0,sK2) | ~r1_ordinal1(sK2,sK0) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
528. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~47) [resolution 479,164] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
530. ~v3_ordinal1(sK0) <- (5, ~47) [subsumption resolution 528,201] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
531. $false <- (5, 6, ~47) [subsumption resolution 530,206] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
532. ~5 | ~6 | 47 [avatar contradiction clause 531]
534. v1_ordinal1(k11_ordinal2(sK0,sK1)) <- (47) [resolution 478,139] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
541. 56 <=> v1_ordinal1(k11_ordinal2(sK0,sK1)) [avatar definition]
543. v1_ordinal1(k11_ordinal2(sK0,sK1)) <- (56) [avatar component clause 541] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
544. 56 | ~47 [avatar split clause 534,477,541]
546. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK2) <- (~48) [resolution 483,164] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
548. ~v3_ordinal1(sK2) <- (5, ~48) [subsumption resolution 546,201] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
549. $false <- (4, 5, ~48) [subsumption resolution 548,196] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
550. ~4 | ~5 | 48 [avatar contradiction clause 549]
552. ~r2_hidden(sK0,sK2) <- (53) [resolution 519,162] {a:5,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
556. 57 <=> r2_hidden(sK0,sK2) [avatar definition]
558. ~r2_hidden(sK0,sK2) <- (~57) [avatar component clause 556] {a:5,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
559. ~57 | ~53 [avatar split clause 552,517,556]
585. 62 <=> r2_xboole_0(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) [avatar definition]
587. r2_xboole_0(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) <- (62) [avatar component clause 585] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
589. 63 <=> k11_ordinal2(sK0,sK1) = k11_ordinal2(sK2,sK1) [avatar definition]
591. k11_ordinal2(sK0,sK1) = k11_ordinal2(sK2,sK1) <- (63) [avatar component clause 589] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
701. 73 <=> r2_hidden(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) [avatar definition]
703. r2_hidden(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) <- (73) [avatar component clause 701] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
733. $false <- (~2, 4, 5, 6, ~52, 63) [unit resulting resolution 196,201,206,507,186,591,141] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:30,thDist:-30}
742. 2 | ~4 | ~5 | ~6 | 52 | ~63 [avatar contradiction clause 733]
753. r1_tarski(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) | ~v3_ordinal1(k11_ordinal2(sK2,sK1)) | ~v3_ordinal1(k11_ordinal2(sK0,sK1)) <- (3) [resolution 191,166] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
754. r1_tarski(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) | ~v3_ordinal1(k11_ordinal2(sK2,sK1)) <- (3, 47) [subsumption resolution 753,478] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
755. ~48 | 49 | ~3 | ~47 [avatar split clause 754,477,189,485,481]
778. k11_ordinal2(sK0,sK1) = k11_ordinal2(sK2,sK1) | r2_xboole_0(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) <- (49) [resolution 487,151] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
781. 62 | 63 | ~49 [avatar split clause 778,485,589,585]
783. r2_hidden(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) | ~v3_ordinal1(k11_ordinal2(sK2,sK1)) | ~v1_ordinal1(k11_ordinal2(sK0,sK1)) <- (62) [resolution 587,133] {a:3,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
785. r2_hidden(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) | ~v1_ordinal1(k11_ordinal2(sK0,sK1)) <- (48, 62) [subsumption resolution 783,482] {a:3,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
786. r2_hidden(k11_ordinal2(sK0,sK1),k11_ordinal2(sK2,sK1)) <- (48, 56, 62) [subsumption resolution 785,543] {a:3,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
787. 73 | ~48 | ~56 | ~62 [avatar split clause 786,585,541,481,701]
838. r2_hidden(sK0,sK2) | ~v3_ordinal1(sK2) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (73) [resolution 703,142] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
844. ~v3_ordinal1(sK2) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~57, 73) [subsumption resolution 838,558] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
845. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (4, ~57, 73) [subsumption resolution 844,196] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
846. ~v3_ordinal1(sK0) <- (4, 5, ~57, 73) [subsumption resolution 845,201] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
847. $false <- (4, 5, 6, ~57, 73) [subsumption resolution 846,206] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
848. ~4 | ~5 | ~6 | 57 | ~73 [avatar contradiction clause 847]
849. $false [avatar sat refutation 182,187,192,197,202,207,350,470,496,509,520,526,532,544,550,559,742,755,781,787,848] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:812,thDist:-812}
% SZS output end Proof for ordinal3__t35_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 45
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 55
% Generated clauses: 394
% Active clauses: 149
% Passive clauses: 172
% Final active clauses: 131
% Final passive clauses: 20
% Main loop iterations started: 152
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 33
% Bw demodulations: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Equational tautologies: 2
% Forward subsumptions: 72
% 
% >>> Generating Inferences
% Binary resolution: 176
% Unit resulting resolution: 5
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 8
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 13
% Split components: 27
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 99
% SAT solver unit clauses: 19
% SAT solver binary clauses: 49
% 
% Memory used [KB]: 5373
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.002 s
% other: 0.005 s ( own 0.001 s ) 

