i: [1,174,0,2,1,1,0,0] t5_tdlat_3
a: [3,271,0,2,1,174]
s: 271
i: [1,173,0,2,1,1,0,0] t5_tdlat_3
a: [3,276,0,2,1,173]
s: 276
i: [1,178,0,13,3,1,0,0] t5_tdlat_3
a: [3,251,0,3,1,178]
s: 251
i: [1,175,0,5,1,1,0,0] t5_tdlat_3
a: [3,266,0,5,1,175]
s: 266
a: [3,258,0,7,1,178]
s: 258
i: [1,179,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,281,0,2,1,179]
s: 281
i: [1,180,0,2,1,0,0,255] fc1_xboole_0
a: [3,286,0,2,1,180]
s: 286
i: [1,243,0,2,1,0,0,2] rc2_xboole_0
a: [3,310,0,2,1,243]
s: 310
i: [1,244,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,315,0,2,1,244]
s: 315
i: [1,245,0,2,1,0,0,2] rc1_xboole_0
a: [3,320,0,2,1,245]
s: 320
i: [1,182,0,3,1,0,0,2] fc1_subset_1
s: 182
i: [1,231,0,3,1,0,0,2] rc2_subset_1
s: 231
i: [1,234,0,3,1,0,0,255] reflexivity_r1_tarski
s: 234
i: [1,246,0,3,1,0,0,255] d10_xboole_0
s: 246
i: [1,247,0,3,1,0,0,255] d10_xboole_0
s: 247
i: [1,181,0,3,1,0,0,255] d2_xboole_0
a: [3,291,0,3,1,181]
s: 291
r: 281
i: [1,229,0,4,1,0,0,1] existence_m1_subset_1
s: 229
i: [1,233,0,4,1,0,0,255] rc3_subset_1
s: 233
i: [1,186,0,5,2,0,0,2] rc1_subset_1
s: 186
i: [1,196,0,5,2,0,0,255] t6_boole
s: 196
i: [1,226,0,5,2,0,0,255] t7_boole
s: 226
i: [1,230,0,5,1,0,0,2] rc2_subset_1
s: 230
i: [1,232,0,5,1,0,0,255] rc3_subset_1
s: 232
i: [1,188,0,6,2,0,0,255] rc4_subset_1
s: 188
i: [1,195,0,6,2,0,0,255] rc5_tops_1
s: 195
i: [1,218,0,6,2,0,0,255] t1_subset
s: 218
i: [1,235,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 235
i: [1,185,0,7,2,0,0,2] rc1_subset_1
s: 185
i: [1,187,0,7,2,0,0,255] rc4_subset_1
s: 187
i: [1,223,0,7,2,0,0,255] t3_subset
s: 223
i: [1,224,0,7,2,0,0,255] t3_subset
s: 224
d: [2,325,1,4,1,1,40,223,266]
a: [3,334,1,4,1,325]
s: 334
i: [1,225,0,7,3,0,0,1] t8_boole
s: 225
d: [2,329,1,4,1,0,40,223,230]
s: 329
i: [1,200,0,24,7,0,0,1] t23_tops_1
a: [3,295,0,7,2,200]
s: 295
d: [2,337,1,0,0,3,52,271,266,295]
e: 337
d: [2,338,1,0,0,2,52,271,229,295]
e: 338
d: [2,339,1,0,0,2,52,271,230,295]
e: 339
d: [2,340,1,0,0,2,52,271,232,295]
e: 340
d: [2,342,1,2,1,2,40,295,266]
d: [2,352,1,0,0,3,21,342,271]
e: 352
r: 295
i: [1,194,0,8,2,0,0,255] rc5_tops_1
s: 194
d: [2,330,1,4,1,0,40,223,232]
s: 330
i: [1,197,0,8,3,0,0,2] cc1_subset_1
s: 197
d: [2,356,1,5,2,1,40,197,266]
a: [3,366,1,3,1,356]
s: 366
i: [1,208,0,8,3,0,0,255] rc7_tops_1
s: 208
i: [1,210,0,8,3,0,0,1] rc6_pre_topc
s: 210
d: [2,328,1,5,1,0,40,223,229]
s: 328
i: [1,212,0,8,3,0,0,1] rc1_tops_1
s: 212
d: [2,362,1,5,2,0,40,197,232]
s: 362
i: [1,214,0,8,3,0,0,1] rc10_pre_topc
s: 214
i: [1,216,0,8,3,0,0,1] rc2_tops_1
s: 216
d: [2,326,1,6,2,0,40,223,185]
s: 326
i: [1,217,0,8,3,0,0,1] rc2_tops_1
s: 217
d: [2,327,1,6,2,0,40,223,187]
s: 327
i: [1,219,0,8,3,0,0,255] t2_subset
s: 219
i: [1,198,0,9,3,0,0,255] cc4_subset_1
s: 198
d: [2,373,1,9,2,1,40,219,266]
d: [2,381,1,5,1,1,21,373,182]
a: [3,385,1,5,1,381]
s: 385
i: [1,222,0,9,3,0,0,255] d10_xboole_0
s: 222
d: [2,409,2,8,2,1,40,222,334]
a: [3,418,2,4,1,409]
s: 418
i: [1,228,0,9,3,0,0,255] t5_subset
s: 228
i: [1,199,0,24,7,0,0,1] t23_tops_1
a: [3,302,0,9,3,199]
s: 302
d: [2,432,1,0,0,4,52,271,276,266,302]
e: 432
d: [2,433,1,0,0,3,52,271,276,229,302]
e: 433
d: [2,434,1,0,0,3,52,271,276,230,302]
e: 434
d: [2,435,1,0,0,3,52,271,276,232,302]
e: 435
d: [2,437,1,4,2,2,40,302,266]
d: [2,449,1,2,1,3,21,437,276]
d: [2,450,1,0,0,4,21,449,271]
e: 450
r: 302
d: [2,400,2,5,1,1,40,385,235]
a: [3,406,2,5,1,400]
s: 406
i: [1,207,0,10,3,0,0,255] rc7_tops_1
s: 207
d: [2,378,1,8,2,0,40,219,230]
d: [2,390,1,5,1,0,21,378,182]
s: 390
i: [1,209,0,10,3,0,0,1] rc6_pre_topc
s: 209
i: [1,211,0,10,3,0,0,1] rc1_tops_1
s: 211
d: [2,379,1,8,2,0,40,219,232]
d: [2,391,1,5,1,0,21,379,182]
s: 391
i: [1,213,0,10,3,0,0,1] rc10_pre_topc
s: 213
d: [2,458,2,5,1,0,40,390,235]
s: 458
i: [1,215,0,10,3,0,0,1] rc2_tops_1
s: 215
i: [1,227,0,10,3,0,0,255] t4_subset
s: 227
d: [2,489,1,7,2,1,40,227,266]
s: 489
i: [1,183,0,11,4,0,0,255] cc3_subset_1
s: 183
a: [3,369,1,2,1,356]
s: 369
i: [1,184,0,11,4,0,0,255] cc2_subset_1
s: 184
i: [1,189,0,12,3,0,0,255] t18_pre_topc
s: 189
d: [2,527,1,7,2,1,40,189,266]
d: [2,545,1,5,1,2,21,527,271]
a: [3,549,1,5,1,545]
s: 549
i: [1,190,0,12,3,0,0,255] t16_tops_1
s: 190
d: [2,551,2,10,2,1,40,549,222]
a: [3,554,2,5,1,551]
s: 554
i: [1,193,0,12,4,0,0,255] cc2_tops_1
s: 193
i: [1,201,0,14,5,0,0,255] cc3_tops_1
s: 201
a: [3,559,2,5,1,551]
s: 559
r: 554
r: 549
r: 258
r: 251
r: 559
a: [3,250,0,3,1,178]
s: 250
a: [3,254,0,3,1,178]
s: 254
a: [3,259,0,7,1,178]
s: 259
i: [1,202,0,14,5,0,0,1] cc2_pre_topc
s: 202
s: 549
i: [1,203,0,14,5,0,0,1] cc1_tops_1
s: 203
d: [2,562,1,7,2,1,40,190,266]
d: [2,580,1,5,1,2,21,562,271]
a: [3,637,1,5,1,580]
s: 637
i: [1,236,0,14,4,0,0,1] fc1_tops_1
s: 236
i: [1,237,0,14,4,0,0,1] fc9_tops_1
s: 237
d: [2,688,1,9,3,1,40,236,266]
d: [2,706,1,7,2,2,21,688,271]
d: [2,707,1,5,1,3,21,706,276]
a: [3,711,1,5,1,707]
s: 711
i: [1,238,0,14,4,0,0,255] fc13_tops_1
s: 238
d: [2,740,1,9,3,1,40,238,266]
d: [2,758,1,7,2,2,21,740,271]
a: [3,762,1,3,1,758]
s: 762
i: [1,239,0,14,3,0,0,1] dt_k2_pre_topc
s: 239
i: [1,240,0,14,3,0,0,1] dt_k1_tops_1
s: 240
d: [2,714,1,9,3,1,40,237,266]
d: [2,732,1,7,2,2,21,714,271]
d: [2,733,1,5,1,3,21,732,276]
a: [3,737,1,5,1,733]
s: 737
i: [1,191,0,15,4,0,0,1] t22_pre_topc
s: 191
s: 559
r: 711
r: 549
r: 259
d: [2,628,0,5,1,2,23,258,559]
a: [3,631,0,5,1,628]
s: 631
i: [1,204,0,15,5,0,0,255] cc4_tops_1
s: 204
d: [2,858,1,10,4,1,40,204,266]
d: [2,878,1,7,3,2,21,858,762]
d: [2,879,1,5,2,3,21,878,271]
d: [2,880,1,3,1,4,21,879,276]
a: [3,884,1,3,1,880]
s: 884
a: [3,305,0,15,4,199]
s: 305
d: [2,887,1,0,0,5,52,271,250,631,266,305]
e: 887
d: [2,891,1,10,3,2,40,305,266]
d: [2,912,1,8,2,3,21,891,271]
d: [2,913,1,3,1,4,21,912,631]
d: [2,914,1,0,0,5,21,913,250]
e: 914
i: [1,177,0,10,2,1,0,0] t5_tdlat_3
i: [1,176,0,10,2,1,0,0] t5_tdlat_3
d: [2,581,1,5,1,3,22,580,258]
d: [2,659,2,10,2,1,40,549,222]
d: [2,824,1,10,3,1,40,191,266]
d: [2,844,1,7,2,2,21,824,254]
d: [2,845,1,5,1,3,21,844,271]
d: [2,852,0,5,1,2,23,259,559]
f: 914,852,845,733,707,659,581,545,435,199,173,174,175,176,177,178
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tdlat_3__t5_tdlat_3
% SZS output start Proof for tdlat_3__t5_tdlat_3
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ((v4_pre_topc(X1,X0) & v3_pre_topc(X1,X0)) <=> k2_pre_topc(X0,X1) = k1_tops_1(X0,X1)))) [input t5_tdlat_3]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ((v4_pre_topc(X1,X0) & v3_pre_topc(X1,X0)) <=> k2_pre_topc(X0,X1) = k1_tops_1(X0,X1)))) [negated conjecture 1]
15. ! [X0] : ! [X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [input d10_xboole_0]
31. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) & l1_pre_topc(X0) & v2_pre_topc(X0)) => v4_pre_topc(k2_pre_topc(X0,X1),X0)) [input fc1_tops_1]
33. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) & l1_pre_topc(X0) & v2_pre_topc(X0)) => v3_pre_topc(k1_tops_1(X0,X1),X0)) [input fc9_tops_1]
43. ! [X0] : ? [X1] : (~v1_subset_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(X0))) [input rc3_subset_1]
49. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => r1_tarski(k1_tops_1(X0,X1),X1))) [input t16_tops_1]
50. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => r1_tarski(X1,k2_pre_topc(X0,X1)))) [input t18_pre_topc]
52. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (((k2_pre_topc(X0,X1) = X1 & v2_pre_topc(X0)) => v4_pre_topc(X1,X0)) & (v4_pre_topc(X1,X0) => k2_pre_topc(X0,X1) = X1)))) [input t22_pre_topc]
53. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (l1_pre_topc(X1) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) => ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))) => ((k1_tops_1(X0,X2) = X2 => v3_pre_topc(X2,X0)) & (v3_pre_topc(X3,X1) => k1_tops_1(X1,X3) = X3)))))) [input t23_tops_1]
63. ! [X0,X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [flattening 15]
72. ? [X0] : (? [X1] : (((v4_pre_topc(X1,X0) & v3_pre_topc(X1,X0)) <~> k2_pre_topc(X0,X1) = k1_tops_1(X0,X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & (l1_pre_topc(X0) & v2_pre_topc(X0))) [ennf transformation 2]
73. ? [X0] : (? [X1] : (((v4_pre_topc(X1,X0) & v3_pre_topc(X1,X0)) <~> k2_pre_topc(X0,X1) = k1_tops_1(X0,X1)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 72]
80. ! [X0] : (! [X1] : (r1_tarski(X1,k2_pre_topc(X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 50]
81. ! [X0] : (! [X1] : (r1_tarski(k1_tops_1(X0,X1),X1) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 49]
82. ! [X0] : (! [X1] : (((v4_pre_topc(X1,X0) | (k2_pre_topc(X0,X1) != X1 | ~v2_pre_topc(X0))) & (k2_pre_topc(X0,X1) = X1 | ~v4_pre_topc(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 52]
83. ! [X0] : (! [X1] : (((v4_pre_topc(X1,X0) | k2_pre_topc(X0,X1) != X1 | ~v2_pre_topc(X0)) & (k2_pre_topc(X0,X1) = X1 | ~v4_pre_topc(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [flattening 82]
90. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((v3_pre_topc(X2,X0) | k1_tops_1(X0,X2) != X2) & (k1_tops_1(X1,X3) = X3 | ~v3_pre_topc(X3,X1))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))) | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X1)) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 53]
91. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((v3_pre_topc(X2,X0) | k1_tops_1(X0,X2) != X2) & (k1_tops_1(X1,X3) = X3 | ~v3_pre_topc(X3,X1))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))) | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X1)) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 90]
123. ! [X0,X1] : (v4_pre_topc(k2_pre_topc(X0,X1),X0) | (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 31]
124. ! [X0,X1] : (v4_pre_topc(k2_pre_topc(X0,X1),X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 123]
125. ! [X0,X1] : (v3_pre_topc(k1_tops_1(X0,X1),X0) | (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 33]
126. ! [X0,X1] : (v3_pre_topc(k1_tops_1(X0,X1),X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 125]
137. ? [X0] : (? [X1] : (((k2_pre_topc(X0,X1) != k1_tops_1(X0,X1) | (~v4_pre_topc(X1,X0) | ~v3_pre_topc(X1,X0))) & (k2_pre_topc(X0,X1) = k1_tops_1(X0,X1) | (v4_pre_topc(X1,X0) & v3_pre_topc(X1,X0)))) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) [nnf transformation 73]
138. ? [X0] : (? [X1] : ((k2_pre_topc(X0,X1) != k1_tops_1(X0,X1) | ~v4_pre_topc(X1,X0) | ~v3_pre_topc(X1,X0)) & (k2_pre_topc(X0,X1) = k1_tops_1(X0,X1) | (v4_pre_topc(X1,X0) & v3_pre_topc(X1,X0))) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) [flattening 137]
139. ? [X0] : (? [X1] : ((k2_pre_topc(X0,X1) != k1_tops_1(X0,X1) | ~v4_pre_topc(X1,X0) | ~v3_pre_topc(X1,X0)) & (k2_pre_topc(X0,X1) = k1_tops_1(X0,X1) | (v4_pre_topc(X1,X0) & v3_pre_topc(X1,X0))) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0)) => (? [X1] : ((k2_pre_topc(sK0,X1) != k1_tops_1(sK0,X1) | ~v4_pre_topc(X1,sK0) | ~v3_pre_topc(X1,sK0)) & (k2_pre_topc(sK0,X1) = k1_tops_1(sK0,X1) | (v4_pre_topc(X1,sK0) & v3_pre_topc(X1,sK0))) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0)) [choice axiom]
140. ? [X1] : ((k2_pre_topc(sK0,X1) != k1_tops_1(sK0,X1) | ~v4_pre_topc(X1,sK0) | ~v3_pre_topc(X1,sK0)) & (k2_pre_topc(sK0,X1) = k1_tops_1(sK0,X1) | (v4_pre_topc(X1,sK0) & v3_pre_topc(X1,sK0))) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => ((k2_pre_topc(sK0,sK1) != k1_tops_1(sK0,sK1) | ~v4_pre_topc(sK1,sK0) | ~v3_pre_topc(sK1,sK0)) & (k2_pre_topc(sK0,sK1) = k1_tops_1(sK0,sK1) | (v4_pre_topc(sK1,sK0) & v3_pre_topc(sK1,sK0))) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
141. ((k2_pre_topc(sK0,sK1) != k1_tops_1(sK0,sK1) | ~v4_pre_topc(sK1,sK0) | ~v3_pre_topc(sK1,sK0)) & (k2_pre_topc(sK0,sK1) = k1_tops_1(sK0,sK1) | (v4_pre_topc(sK1,sK0) & v3_pre_topc(sK1,sK0))) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) [skolemisation 138,140,139]
158. ! [X0,X1] : ((X0 = X1 | (~r1_tarski(X1,X0) | ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [nnf transformation 63]
159. ! [X0,X1] : ((X0 = X1 | ~r1_tarski(X1,X0) | ~r1_tarski(X0,X1)) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [flattening 158]
165. ! [X0] : (? [X1] : (~v1_subset_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (~v1_subset_1(sK12(X0),X0) & m1_subset_1(sK12(X0),k1_zfmisc_1(X0)))) [choice axiom]
166. ! [X0] : (~v1_subset_1(sK12(X0),X0) & m1_subset_1(sK12(X0),k1_zfmisc_1(X0))) [skolemisation 43,165]
173. v2_pre_topc(sK0) [cnf transformation 141] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
174. l1_pre_topc(sK0) [cnf transformation 141] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
175. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 141] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
176. k2_pre_topc(sK0,sK1) = k1_tops_1(sK0,sK1) | v3_pre_topc(sK1,sK0) [cnf transformation 141] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
177. k2_pre_topc(sK0,sK1) = k1_tops_1(sK0,sK1) | v4_pre_topc(sK1,sK0) [cnf transformation 141] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
178. k2_pre_topc(sK0,sK1) != k1_tops_1(sK0,sK1) | ~v4_pre_topc(sK1,sK0) | ~v3_pre_topc(sK1,sK0) [cnf transformation 141] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
189. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | r1_tarski(X1,k2_pre_topc(X0,X1)) | ~l1_pre_topc(X0) [cnf transformation 80] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
190. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | r1_tarski(k1_tops_1(X0,X1),X1) | ~l1_pre_topc(X0) [cnf transformation 81] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
191. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v4_pre_topc(X1,X0) | k2_pre_topc(X0,X1) = X1 | ~l1_pre_topc(X0) [cnf transformation 83] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
199. k1_tops_1(X1,X3) = X3 | ~v3_pre_topc(X3,X1) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))) | ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X1) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 91] {a:0,w:24,wCS:9600,sine:1,thAx:0,allAx:1,thDist:-1}
222. ~r1_tarski(X1,X0) | X0 = X1 | ~r1_tarski(X0,X1) [cnf transformation 159] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
232. m1_subset_1(sK12(X0),k1_zfmisc_1(X0)) [cnf transformation 166] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
236. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | v4_pre_topc(k2_pre_topc(X0,X1),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 124] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
237. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | v3_pre_topc(k1_tops_1(X0,X1),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 126] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
249. 1 <=> v3_pre_topc(sK1,sK0) [avatar definition]
250. v3_pre_topc(sK1,sK0) <- (1) [avatar component clause 249] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
253. 2 <=> v4_pre_topc(sK1,sK0) [avatar definition]
254. v4_pre_topc(sK1,sK0) <- (2) [avatar component clause 253] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
257. 3 <=> k2_pre_topc(sK0,sK1) = k1_tops_1(sK0,sK1) [avatar definition]
258. k2_pre_topc(sK0,sK1) = k1_tops_1(sK0,sK1) <- (3) [avatar component clause 257] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
259. k2_pre_topc(sK0,sK1) != k1_tops_1(sK0,sK1) <- (~3) [avatar component clause 257] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
260. ~1 | ~2 | ~3 [avatar split clause 178,257,253,249]
261. 2 | 3 [avatar split clause 177,257,253]
262. 1 | 3 [avatar split clause 176,257,249]
264. 4 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
266. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (4) [avatar component clause 264] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
267. 4 [avatar split clause 175,264]
269. 5 <=> l1_pre_topc(sK0) [avatar definition]
271. l1_pre_topc(sK0) <- (5) [avatar component clause 269] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
272. 5 [avatar split clause 174,269]
274. 6 <=> v2_pre_topc(sK0) [avatar definition]
276. v2_pre_topc(sK0) <- (6) [avatar component clause 274] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
277. 6 [avatar split clause 173,274]
301. 12 <=> ! [X0,X2] : (~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) | ~v2_pre_topc(X0) | ~l1_pre_topc(X0)) [avatar definition]
302. ~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X0))) | ~v2_pre_topc(X0) | ~l1_pre_topc(X0) <- (12) [avatar component clause 301] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
304. 13 <=> ! [X1,X3] : (k1_tops_1(X1,X3) = X3 | ~l1_pre_topc(X1) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))) | ~v3_pre_topc(X3,X1)) [avatar definition]
305. ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))) | ~l1_pre_topc(X1) | k1_tops_1(X1,X3) = X3 | ~v3_pre_topc(X3,X1) <- (13) [avatar component clause 304] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
306. 12 | 13 [avatar split clause 199,304,301]
435. $false <- (5, 6, 12) [unit resulting resolution 271,276,232,302] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
448. ~5 | ~6 | ~12 [avatar contradiction clause 435]
527. r1_tarski(sK1,k2_pre_topc(sK0,sK1)) | ~l1_pre_topc(sK0) <- (4) [resolution 189,266] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
545. r1_tarski(sK1,k2_pre_topc(sK0,sK1)) <- (4, 5) [subsumption resolution 527,271] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
547. 25 <=> r1_tarski(sK1,k2_pre_topc(sK0,sK1)) [avatar definition]
549. r1_tarski(sK1,k2_pre_topc(sK0,sK1)) <- (25) [avatar component clause 547] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
550. 25 | ~4 | ~5 [avatar split clause 545,269,264,547]
553. 26 <=> r1_tarski(k2_pre_topc(sK0,sK1),sK1) [avatar definition]
557. 27 <=> sK1 = k2_pre_topc(sK0,sK1) [avatar definition]
559. sK1 = k2_pre_topc(sK0,sK1) <- (27) [avatar component clause 557] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
562. r1_tarski(k1_tops_1(sK0,sK1),sK1) | ~l1_pre_topc(sK0) <- (4) [resolution 190,266] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
580. r1_tarski(k1_tops_1(sK0,sK1),sK1) <- (4, 5) [subsumption resolution 562,271] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
581. r1_tarski(k2_pre_topc(sK0,sK1),sK1) <- (3, 4, 5) [forward demodulation 580,258] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
582. 26 | ~3 | ~4 | ~5 [avatar split clause 581,269,264,257,553]
630. 28 <=> sK1 = k1_tops_1(sK0,sK1) [avatar definition]
631. sK1 != k1_tops_1(sK0,sK1) <- (~28) [avatar component clause 630] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
659. sK1 = k2_pre_topc(sK0,sK1) | ~r1_tarski(k2_pre_topc(sK0,sK1),sK1) <- (25) [resolution 549,222] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
660. ~26 | 27 | ~25 [avatar split clause 659,547,557,553]
688. v4_pre_topc(k2_pre_topc(sK0,sK1),sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (4) [resolution 236,266] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
706. v4_pre_topc(k2_pre_topc(sK0,sK1),sK0) | ~v2_pre_topc(sK0) <- (4, 5) [subsumption resolution 688,271] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
707. v4_pre_topc(k2_pre_topc(sK0,sK1),sK0) <- (4, 5, 6) [subsumption resolution 706,276] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
709. 31 <=> v4_pre_topc(k2_pre_topc(sK0,sK1),sK0) [avatar definition]
712. 31 | ~4 | ~5 | ~6 [avatar split clause 707,274,269,264,709]
714. v3_pre_topc(k1_tops_1(sK0,sK1),sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (4) [resolution 237,266] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
732. v3_pre_topc(k1_tops_1(sK0,sK1),sK0) | ~v2_pre_topc(sK0) <- (4, 5) [subsumption resolution 714,271] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
733. v3_pre_topc(k1_tops_1(sK0,sK1),sK0) <- (4, 5, 6) [subsumption resolution 732,276] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
735. 32 <=> v3_pre_topc(k1_tops_1(sK0,sK1),sK0) [avatar definition]
738. 32 | ~4 | ~5 | ~6 [avatar split clause 733,274,269,264,735]
824. ~v4_pre_topc(sK1,sK0) | sK1 = k2_pre_topc(sK0,sK1) | ~l1_pre_topc(sK0) <- (4) [resolution 191,266] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
844. sK1 = k2_pre_topc(sK0,sK1) | ~l1_pre_topc(sK0) <- (2, 4) [subsumption resolution 824,254] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
845. sK1 = k2_pre_topc(sK0,sK1) <- (2, 4, 5) [subsumption resolution 844,271] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
846. 27 | ~2 | ~4 | ~5 [avatar split clause 845,269,264,253,557]
852. sK1 != k1_tops_1(sK0,sK1) <- (~3, 27) [backward demodulation 259,559] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
854. ~28 | 3 | ~27 [avatar split clause 852,557,257,630]
891. ~l1_pre_topc(sK0) | sK1 = k1_tops_1(sK0,sK1) | ~v3_pre_topc(sK1,sK0) <- (4, 13) [resolution 305,266] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
912. sK1 = k1_tops_1(sK0,sK1) | ~v3_pre_topc(sK1,sK0) <- (4, 5, 13) [subsumption resolution 891,271] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
913. ~v3_pre_topc(sK1,sK0) <- (4, 5, 13, ~28) [subsumption resolution 912,631] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
914. $false <- (1, 4, 5, 13, ~28) [subsumption resolution 913,250] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
915. ~1 | ~4 | ~5 | ~13 | 28 [avatar contradiction clause 914]
920. sK1 != k2_pre_topc(sK0,sK1) | v4_pre_topc(sK1,sK0) | ~v4_pre_topc(k2_pre_topc(sK0,sK1),sK0) [theory tautology sat conflict] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:1,thDist:-1}
921. sK1 != k2_pre_topc(sK0,sK1) | k2_pre_topc(sK0,sK1) != k1_tops_1(sK0,sK1) | v3_pre_topc(sK1,sK0) | ~v3_pre_topc(k1_tops_1(sK0,sK1),sK0) [theory tautology sat conflict] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:1,thDist:-1}
922. $false [avatar sat refutation 260,261,262,267,272,277,306,448,550,582,660,712,738,846,854,915,920,921] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:290,thDist:-290}
% SZS output end Proof for tdlat_3__t5_tdlat_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 59
% >>> Preprocessing
% Introduced skolems: 16
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 73
% Generated clauses: 600
% Active clauses: 95
% Passive clauses: 334
% Final active clauses: 84
% Final passive clauses: 234
% Main loop iterations started: 98
% 
% >>> Simplifying Inferences
% Duplicate literals: 129
% Fw subsumption resolutions: 48
% Fw demodulations: 1
% Bw demodulations: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Equational tautologies: 1
% Forward subsumptions: 38
% 
% >>> Generating Inferences
% Binary resolution: 321
% Unit resulting resolution: 9
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 25
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 52
% SAT solver unit clauses: 9
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 5373
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% binary resolution index maintenance: 0.001 s
% resolution: 0.001 s
% other: 0.007 s ( own 0.003 s ) 

