i: [1,37,0,15,3,1,0,0] t107_xboole_1
a: [3,53,0,5,1,37]
s: 53
a: [3,56,0,5,1,37]
s: 56
a: [3,60,0,5,1,37]
s: 60
i: [1,43,0,3,1,0,0,1] reflexivity_r1_tarski
s: 43
i: [1,44,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 44
i: [1,45,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 45
i: [1,49,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 49
r: 53
r: 60
a: [3,52,0,5,1,37]
s: 52
a: [3,61,0,5,1,37]
s: 61
i: [1,46,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 46
i: [1,47,0,7,1,0,0,1] commutativity_k5_xboole_0
s: 47
i: [1,48,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 48
i: [1,40,0,8,2,0,0,1] t106_xboole_1
s: 40
d: [2,71,1,5,1,0,40,40,43]
s: 71
i: [1,41,0,8,2,0,0,1] t106_xboole_1
s: 41
i: [1,38,0,11,1,0,0,1] l98_xboole_1
s: 38
d: [2,73,1,5,1,0,40,41,43]
s: 73
i: [1,39,0,11,1,0,0,1] d6_xboole_0
s: 39
d: [2,86,2,5,1,0,40,73,49]
s: 86
i: [1,42,0,11,3,0,0,1] t86_xboole_1
s: 42
d: [2,75,1,9,1,0,44,38,45]
d: [2,84,1,7,1,0,22,75,44]
s: 84
d: [2,108,2,5,1,0,44,71,84]
s: 108
d: [2,78,1,9,1,0,44,38,44]
d: [2,85,1,7,1,0,22,78,45]
s: 85
d: [2,109,2,5,1,0,44,73,84]
s: 109
d: [2,90,1,7,1,0,44,39,45]
s: 90
d: [2,93,1,7,1,0,44,45,39]
s: 93
d: [2,110,3,5,1,0,44,86,84]
s: 110
d: [2,81,1,10,2,0,44,41,38]
s: 81
d: [2,115,2,0,0,2,52,61,52,81]
e: 115
d: [2,119,2,5,1,1,40,81,52]
d: [2,123,2,0,0,2,21,119,61]
e: 123
r: 56
r: 61
a: [3,57,0,5,1,37]
s: 57
s: 60
d: [2,65,1,5,1,1,40,49,60]
a: [3,69,1,5,1,65]
s: 69
d: [2,113,3,5,1,0,40,109,49]
s: 113
d: [2,83,1,10,2,0,44,40,38]
s: 83
d: [2,129,2,0,0,2,52,57,52,83]
e: 129
d: [2,133,2,5,1,1,40,83,52]
d: [2,137,2,0,0,2,21,133,57]
e: 137
r: 52
r: 57
s: 53
s: 56
d: [2,76,1,11,1,0,44,38,48]
s: 76
d: [2,72,2,7,1,0,40,71,40]
s: 72
d: [2,77,1,11,1,0,44,38,48]
s: 77
d: [2,74,2,7,1,0,40,41,71]
s: 74
d: [2,79,1,11,1,0,44,38,46]
s: 79
d: [2,80,1,11,1,0,44,38,46]
s: 80
d: [2,82,2,7,1,0,44,71,38]
s: 82
d: [2,91,1,11,1,0,44,39,48]
s: 91
d: [2,87,2,7,1,0,44,73,38]
s: 87
d: [2,92,1,11,1,0,44,39,48]
s: 92
d: [2,95,1,11,1,0,44,48,39]
s: 95
d: [2,116,2,7,1,0,40,81,43]
s: 116
d: [2,96,1,11,1,0,44,48,39]
s: 96
d: [2,130,2,7,1,0,40,83,43]
s: 130
d: [2,102,1,15,3,0,44,42,38]
s: 102
d: [2,227,2,0,0,3,52,56,53,60,102]
e: 227
d: [2,232,2,10,2,1,40,102,60]
d: [2,237,2,5,1,2,21,232,53]
d: [2,238,2,0,0,3,21,237,56]
e: 238
i: [1,36,0,10,2,1,0,0] t107_xboole_1
i: [1,35,0,10,2,1,0,0] t107_xboole_1
f: 238,137,119,35,36,37
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xboole_1__t107_xboole_1
% SZS output start Proof for xboole_1__t107_xboole_1
1. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,k5_xboole_0(X1,X2)) <=> (r1_xboole_0(X0,k3_xboole_0(X1,X2)) & r1_tarski(X0,k2_xboole_0(X1,X2)))) [input t107_xboole_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,k5_xboole_0(X1,X2)) <=> (r1_xboole_0(X0,k3_xboole_0(X1,X2)) & r1_tarski(X0,k2_xboole_0(X1,X2)))) [negated conjecture 1]
13. ! [X0] : ! [X1] : k5_xboole_0(X0,X1) = k4_xboole_0(k2_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [input l98_xboole_1]
16. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,k4_xboole_0(X1,X2)) => (r1_xboole_0(X0,X2) & r1_tarski(X0,X1))) [input t106_xboole_1]
17. ! [X0] : ! [X1] : ! [X2] : ((r1_xboole_0(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k4_xboole_0(X1,X2))) [input t86_xboole_1]
18. ~! [X0,X1,X2] : (r1_tarski(X0,k5_xboole_0(X1,X2)) <=> (r1_xboole_0(X0,k3_xboole_0(X1,X2)) & r1_tarski(X0,k2_xboole_0(X1,X2)))) [flattening 2]
19. ! [X0,X1] : k5_xboole_0(X0,X1) = k4_xboole_0(k2_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [flattening 13]
21. ! [X0,X1,X2] : (r1_tarski(X0,k4_xboole_0(X1,X2)) => (r1_xboole_0(X0,X2) & r1_tarski(X0,X1))) [flattening 16]
22. ! [X0,X1,X2] : ((r1_xboole_0(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k4_xboole_0(X1,X2))) [flattening 17]
26. ? [X0,X1,X2] : (r1_tarski(X0,k5_xboole_0(X1,X2)) <~> (r1_xboole_0(X0,k3_xboole_0(X1,X2)) & r1_tarski(X0,k2_xboole_0(X1,X2)))) [ennf transformation 18]
27. ! [X0,X1,X2] : ((r1_xboole_0(X0,X2) & r1_tarski(X0,X1)) | ~r1_tarski(X0,k4_xboole_0(X1,X2))) [ennf transformation 21]
28. ! [X0,X1,X2] : (r1_tarski(X0,k4_xboole_0(X1,X2)) | (~r1_xboole_0(X0,X2) | ~r1_tarski(X0,X1))) [ennf transformation 22]
29. ! [X0,X1,X2] : (r1_tarski(X0,k4_xboole_0(X1,X2)) | ~r1_xboole_0(X0,X2) | ~r1_tarski(X0,X1)) [flattening 28]
31. ? [X0,X1,X2] : (((~r1_xboole_0(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,k2_xboole_0(X1,X2))) | ~r1_tarski(X0,k5_xboole_0(X1,X2))) & ((r1_xboole_0(X0,k3_xboole_0(X1,X2)) & r1_tarski(X0,k2_xboole_0(X1,X2))) | r1_tarski(X0,k5_xboole_0(X1,X2)))) [nnf transformation 26]
32. ? [X0,X1,X2] : ((~r1_xboole_0(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,k2_xboole_0(X1,X2)) | ~r1_tarski(X0,k5_xboole_0(X1,X2))) & ((r1_xboole_0(X0,k3_xboole_0(X1,X2)) & r1_tarski(X0,k2_xboole_0(X1,X2))) | r1_tarski(X0,k5_xboole_0(X1,X2)))) [flattening 31]
33. ? [X0,X1,X2] : ((~r1_xboole_0(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,k2_xboole_0(X1,X2)) | ~r1_tarski(X0,k5_xboole_0(X1,X2))) & ((r1_xboole_0(X0,k3_xboole_0(X1,X2)) & r1_tarski(X0,k2_xboole_0(X1,X2))) | r1_tarski(X0,k5_xboole_0(X1,X2)))) => ((~r1_xboole_0(sK0,k3_xboole_0(sK1,sK2)) | ~r1_tarski(sK0,k2_xboole_0(sK1,sK2)) | ~r1_tarski(sK0,k5_xboole_0(sK1,sK2))) & ((r1_xboole_0(sK0,k3_xboole_0(sK1,sK2)) & r1_tarski(sK0,k2_xboole_0(sK1,sK2))) | r1_tarski(sK0,k5_xboole_0(sK1,sK2)))) [choice axiom]
34. (~r1_xboole_0(sK0,k3_xboole_0(sK1,sK2)) | ~r1_tarski(sK0,k2_xboole_0(sK1,sK2)) | ~r1_tarski(sK0,k5_xboole_0(sK1,sK2))) & ((r1_xboole_0(sK0,k3_xboole_0(sK1,sK2)) & r1_tarski(sK0,k2_xboole_0(sK1,sK2))) | r1_tarski(sK0,k5_xboole_0(sK1,sK2))) [skolemisation 32,33]
35. r1_tarski(sK0,k2_xboole_0(sK1,sK2)) | r1_tarski(sK0,k5_xboole_0(sK1,sK2)) [cnf transformation 34] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. r1_xboole_0(sK0,k3_xboole_0(sK1,sK2)) | r1_tarski(sK0,k5_xboole_0(sK1,sK2)) [cnf transformation 34] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. ~r1_xboole_0(sK0,k3_xboole_0(sK1,sK2)) | ~r1_tarski(sK0,k2_xboole_0(sK1,sK2)) | ~r1_tarski(sK0,k5_xboole_0(sK1,sK2)) [cnf transformation 34] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. k5_xboole_0(X0,X1) = k4_xboole_0(k2_xboole_0(X0,X1),k3_xboole_0(X0,X1)) [cnf transformation 19] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
40. ~r1_tarski(X0,k4_xboole_0(X1,X2)) | r1_tarski(X0,X1) [cnf transformation 27] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
41. ~r1_tarski(X0,k4_xboole_0(X1,X2)) | r1_xboole_0(X0,X2) [cnf transformation 27] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. r1_tarski(X0,k4_xboole_0(X1,X2)) | ~r1_xboole_0(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 29] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. 1 <=> r1_tarski(sK0,k5_xboole_0(sK1,sK2)) [avatar definition]
52. r1_tarski(sK0,k5_xboole_0(sK1,sK2)) <- (1) [avatar component clause 51] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
53. ~r1_tarski(sK0,k5_xboole_0(sK1,sK2)) <- (~1) [avatar component clause 51] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
55. 2 <=> r1_tarski(sK0,k2_xboole_0(sK1,sK2)) [avatar definition]
56. r1_tarski(sK0,k2_xboole_0(sK1,sK2)) <- (2) [avatar component clause 55] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
57. ~r1_tarski(sK0,k2_xboole_0(sK1,sK2)) <- (~2) [avatar component clause 55] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
59. 3 <=> r1_xboole_0(sK0,k3_xboole_0(sK1,sK2)) [avatar definition]
60. r1_xboole_0(sK0,k3_xboole_0(sK1,sK2)) <- (3) [avatar component clause 59] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. ~1 | ~2 | ~3 [avatar split clause 37,59,55,51]
63. 1 | 3 [avatar split clause 36,59,51]
64. 1 | 2 [avatar split clause 35,55,51]
81. ~r1_tarski(X2,k5_xboole_0(X0,X1)) | r1_xboole_0(X2,k3_xboole_0(X0,X1)) [superposition 41,38] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
83. ~r1_tarski(X7,k5_xboole_0(X5,X6)) | r1_tarski(X7,k2_xboole_0(X5,X6)) [superposition 40,38] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
102. ~r1_xboole_0(X2,k3_xboole_0(X0,X1)) | r1_tarski(X2,k5_xboole_0(X0,X1)) | ~r1_tarski(X2,k2_xboole_0(X0,X1)) [superposition 42,38] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
119. r1_xboole_0(sK0,k3_xboole_0(sK1,sK2)) <- (1) [resolution 81,52] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
125. 3 | ~1 [avatar split clause 119,51,59]
133. r1_tarski(sK0,k2_xboole_0(sK1,sK2)) <- (1) [resolution 83,52] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
137. $false <- (1, ~2) [subsumption resolution 133,57] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
138. ~1 | 2 [avatar contradiction clause 137]
232. r1_tarski(sK0,k5_xboole_0(sK1,sK2)) | ~r1_tarski(sK0,k2_xboole_0(sK1,sK2)) <- (3) [resolution 102,60] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
237. ~r1_tarski(sK0,k2_xboole_0(sK1,sK2)) <- (~1, 3) [subsumption resolution 232,53] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
238. $false <- (~1, 2, 3) [subsumption resolution 237,56] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
239. 1 | ~2 | ~3 [avatar contradiction clause 238]
241. $false [avatar sat refutation 62,63,64,125,138,239] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:61,thDist:-61}
% SZS output end Proof for xboole_1__t107_xboole_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 199
% Active clauses: 39
% Passive clauses: 122
% Final active clauses: 33
% Final passive clauses: 72
% Main loop iterations started: 50
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 4
% Fw demodulations: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 41
% 
% >>> Generating Inferences
% Binary resolution: 34
% Unit resulting resolution: 3
% Forward superposition: 67
% Backward superposition: 42
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver binary clauses: 9
% 
% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% superposition: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

