i: [1,127,0,2,1,1,0,0] t21_mod_2
a: [3,204,0,2,1,127]
s: 204
i: [1,126,0,2,1,1,0,0] t21_mod_2
a: [3,209,0,2,1,126]
s: 209
i: [1,128,0,9,3,1,0,0] t21_mod_2
a: [3,191,0,3,1,128]
s: 191
i: [1,129,0,2,1,0,0,1] spc2_boole
a: [3,214,0,2,1,129]
s: 214
i: [1,130,0,2,1,0,0,1] spc1_boole
a: [3,219,0,2,1,130]
s: 219
i: [1,131,0,2,1,0,0,255] spc0_boole
a: [3,224,0,2,1,131]
s: 224
i: [1,132,0,2,1,0,0,2] fc1_xboole_0
a: [3,229,0,2,1,132]
s: 229
i: [1,184,0,2,1,0,0,1] rc2_xboole_0
a: [3,314,0,2,1,184]
s: 314
i: [1,185,0,2,1,0,0,1] rc1_xboole_0
a: [3,319,0,2,1,185]
s: 319
i: [1,187,0,2,1,0,0,1] rc1_classes2
a: [3,324,0,2,1,187]
s: 324
i: [1,186,0,2,1,0,0,1] rc1_classes2
a: [3,329,0,2,1,186]
s: 329
i: [1,149,0,3,1,0,0,3] fc1_subset_1
s: 149
i: [1,175,0,3,1,0,0,3] rc2_subset_1
s: 175
i: [1,178,0,3,1,0,0,255] reflexivity_r1_tarski
s: 178
i: [1,133,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,234,0,3,1,133]
s: 234
r: 229
d: [2,331,0,2,1,2,23,229,234]
a: [3,335,0,2,1,331]
s: 335
i: [1,134,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,239,0,3,1,134]
s: 239
i: [1,135,0,3,1,0,0,255] t1_numerals
a: [3,244,0,3,1,135]
s: 244
i: [1,142,0,3,1,0,0,255] spc2_numerals
a: [3,269,0,3,1,142]
s: 269
i: [1,141,0,3,1,0,0,255] spc2_numerals
a: [3,274,0,3,1,141]
s: 274
i: [1,145,0,3,1,0,0,255] spc1_numerals
a: [3,284,0,3,1,145]
s: 284
i: [1,144,0,3,1,0,0,255] spc1_numerals
a: [3,289,0,3,1,144]
s: 289
i: [1,148,0,3,1,0,0,255] spc0_numerals
a: [3,299,0,3,1,148]
s: 299
i: [1,147,0,3,1,0,0,255] spc0_numerals
a: [3,304,0,3,1,147]
s: 304
d: [2,337,0,3,1,2,22,244,239]
d: [2,338,0,3,1,3,22,337,234]
a: [3,342,0,3,1,338]
s: 342
i: [1,156,0,4,2,0,0,255] cc2_classes2
s: 156
d: [2,344,1,2,1,1,40,156,204]
a: [3,349,1,2,1,344]
s: 349
i: [1,157,0,4,2,0,0,255] cc2_classes2
s: 157
d: [2,356,1,2,1,1,40,157,204]
a: [3,361,1,2,1,356]
s: 361
i: [1,173,0,4,1,0,0,1] existence_m1_subset_1
s: 173
i: [1,177,0,4,1,0,0,255] rc3_subset_1
s: 177
d: [2,345,1,2,1,1,40,156,324]
a: [3,354,1,2,1,345]
s: 354
i: [1,136,0,4,1,0,0,255] rqSucc__k1_ordinal1__r1_r2
a: [3,249,0,4,1,136]
s: 249
d: [2,357,1,2,1,1,40,157,324]
a: [3,366,1,2,1,357]
s: 366
i: [1,137,0,4,1,0,0,255] rqSucc__k1_ordinal1__r0_r1
a: [3,254,0,4,1,137]
s: 254
i: [1,138,0,4,1,0,0,2] dt_k5_numbers
a: [3,259,0,4,1,138]
s: 259
i: [1,139,0,4,1,0,0,1] dt_k6_numbers
a: [3,264,0,4,1,139]
s: 264
i: [1,140,0,4,1,0,0,255] spc2_numerals
a: [3,279,0,4,1,140]
s: 279
i: [1,143,0,4,1,0,0,255] spc1_numerals
a: [3,294,0,4,1,143]
s: 294
i: [1,146,0,4,1,0,0,255] spc0_numerals
a: [3,309,0,4,1,146]
s: 309
i: [1,153,0,5,2,0,0,3] rc1_subset_1
s: 153
i: [1,158,0,5,2,0,0,2] t6_boole
s: 158
i: [1,170,0,5,2,0,0,3] t7_boole
s: 170
i: [1,174,0,5,1,0,0,3] rc2_subset_1
s: 174
i: [1,176,0,5,1,0,0,255] rc3_subset_1
s: 176
d: [2,368,0,5,2,1,22,158,234]
s: 368
d: [2,373,1,5,2,2,44,191,368]
s: 373
i: [1,155,0,6,2,0,0,255] rc4_subset_1
s: 155
d: [2,374,2,3,1,2,40,373,173]
a: [3,378,2,3,1,374]
s: 378
i: [1,165,0,6,3,0,0,255] cc3_classes2
s: 165
i: [1,166,0,6,2,0,0,3] t1_subset
s: 166
d: [2,371,1,5,2,2,44,342,368]
s: 371
i: [1,179,0,6,2,0,0,3] antisymmetry_r2_hidden
s: 179
d: [2,370,1,6,2,2,44,264,368]
s: 370
i: [1,152,0,7,2,0,0,3] rc1_subset_1
s: 152
i: [1,154,0,7,2,0,0,255] rc4_subset_1
s: 154
i: [1,161,0,7,3,0,0,2] t56_classes2
s: 161
i: [1,164,0,7,3,0,0,1] rc2_classes2
s: 164
i: [1,168,0,7,2,0,0,255] t3_subset
s: 168
i: [1,169,0,7,3,0,0,2] t8_boole
s: 169
d: [2,382,0,7,3,1,22,161,234]
s: 382
d: [2,369,1,7,3,1,44,368,368]
s: 369
i: [1,159,0,8,3,0,0,3] cc1_subset_1
s: 159
d: [2,388,1,4,2,1,40,159,259]
a: [3,397,1,2,1,388]
s: 397
i: [1,163,0,8,3,0,0,1] rc2_classes2
s: 163
i: [1,167,0,8,3,0,0,3] t2_subset
s: 167
r: 397
a: [3,396,1,2,1,388]
s: 396
i: [1,160,0,9,3,0,0,255] cc4_subset_1
s: 160
a: [3,401,1,2,1,388]
s: 401
i: [1,172,0,9,3,0,0,3] t5_subset
s: 172
i: [1,171,0,10,3,0,0,3] t4_subset
s: 171
d: [2,480,1,5,2,1,40,160,259]
d: [2,487,1,3,1,2,21,480,396]
a: [3,491,1,3,1,487]
s: 491
i: [1,150,0,11,4,0,0,255] cc3_subset_1
s: 150
d: [2,495,1,5,2,1,40,172,259]
d: [2,502,1,3,1,2,21,495,396]
s: 502
i: [1,151,0,11,4,0,0,255] cc2_subset_1
s: 151
i: [1,162,0,11,4,0,0,255] t20_mod_2
s: 162
d: [2,526,2,10,4,2,40,162,373]
d: [2,536,2,8,3,3,21,526,204]
d: [2,537,2,6,2,4,21,536,209]
s: 537
i: [1,183,0,14,4,0,0,2] existence_m2_subset_1
s: 183
d: [2,544,3,8,3,4,40,537,163]
d: [2,548,3,6,2,4,21,544,204]
d: [2,549,3,4,1,4,21,548,209]
a: [3,553,3,4,1,549]
s: 553
d: [2,560,0,12,3,0,21,183,159]
s: 560
i: [1,180,0,15,5,0,0,2] dt_m2_subset_1
s: 180
d: [2,545,3,4,1,4,40,537,173]
a: [3,558,3,4,1,545]
s: 558
d: [2,561,0,13,4,0,21,180,159]
s: 561
d: [2,543,3,11,4,4,40,537,162]
d: [2,546,3,9,3,4,21,543,204]
d: [2,547,3,7,2,4,21,546,209]
s: 547
a: [3,195,0,3,1,128]
s: 195
i: [1,181,0,15,5,0,0,2] redefinition_m2_subset_1
s: 181
d: [2,413,1,7,2,1,40,167,259]
d: [2,455,1,4,1,1,21,413,149]
a: [3,459,1,4,1,455]
s: 459
d: [2,577,0,13,4,0,21,181,159]
s: 577
d: [2,393,1,5,2,0,40,159,176]
s: 393
i: [1,182,0,15,5,0,0,2] redefinition_m2_subset_1
s: 182
d: [2,588,0,13,4,0,21,182,159]
s: 588
d: [2,420,1,8,2,0,40,167,174]
d: [2,474,1,5,1,0,21,420,149]
s: 474
d: [2,421,1,8,2,0,40,167,176]
d: [2,475,1,5,1,0,21,421,149]
s: 475
d: [2,520,1,9,3,0,40,150,176]
d: [2,524,1,5,2,0,21,520,177]
s: 524
d: [2,391,1,6,2,0,40,159,173]
s: 391
d: [2,419,1,6,2,0,40,167,173]
s: 419
d: [2,601,2,5,1,0,40,474,179]
s: 601
d: [2,485,1,6,2,0,40,160,174]
s: 485
d: [2,604,2,5,1,0,40,475,179]
s: 604
d: [2,500,1,6,2,0,40,172,174]
s: 500
d: [2,501,1,6,2,0,40,172,176]
s: 501
d: [2,505,1,6,2,1,40,171,259]
s: 505
d: [2,519,1,9,3,0,40,150,174]
d: [2,523,1,6,2,0,21,519,175]
s: 523
d: [2,608,2,6,2,0,40,419,179]
s: 608
d: [2,383,1,7,3,1,40,382,179]
s: 383
d: [2,384,1,7,3,1,40,382,166]
s: 384
d: [2,622,2,0,0,4,52,204,209,191,384]
e: 622
d: [2,623,2,4,2,2,40,384,191]
d: [2,636,2,2,1,3,21,623,209]
d: [2,637,2,0,0,4,21,636,204]
e: 637
d: [2,625,2,6,3,2,40,384,373]
d: [2,641,2,4,2,3,21,625,209]
d: [2,642,2,2,1,4,21,641,204]
d: [2,643,2,0,0,5,21,642,335]
e: 643
r: 547
r: 537
r: 373
r: 191
a: [3,190,0,3,1,128]
s: 190
d: [2,657,1,5,2,1,40,190,167]
d: [2,659,1,3,1,2,21,657,209]
a: [3,663,1,3,1,659]
s: 663
d: [2,665,2,3,1,1,40,663,179]
a: [3,672,2,3,1,665]
s: 672
d: [2,658,1,5,2,2,44,190,368]
s: 658
d: [2,387,1,7,3,0,40,159,168]
s: 387
d: [2,668,2,5,2,2,44,663,368]
s: 668
d: [2,407,1,10,3,0,40,167,168]
d: [2,427,1,7,2,0,21,407,149]
s: 427
d: [2,676,2,7,3,2,40,658,167]
d: [2,677,2,5,2,3,21,676,209]
s: 677
d: [2,416,1,10,3,0,40,167,152]
d: [2,472,1,7,2,0,21,416,149]
s: 472
d: [2,417,1,10,3,0,40,167,154]
d: [2,473,1,7,2,0,21,417,149]
s: 473
d: [2,629,2,11,4,1,40,384,172]
d: [2,648,2,8,3,1,21,629,149]
a: [3,652,2,3,1,648]
s: 652
r: 672
r: 383
d: [2,484,1,7,2,0,40,160,173]
s: 484
d: [2,499,1,7,2,0,40,172,173]
s: 499
d: [2,510,1,7,2,0,40,171,174]
s: 510
d: [2,511,1,7,2,0,40,171,176]
s: 511
d: [2,682,2,7,2,0,40,427,179]
s: 682
d: [2,418,1,10,4,0,40,167,163]
d: [2,422,1,8,3,0,18,418]
s: 422
d: [2,685,2,7,2,0,40,472,179]
s: 685
d: [2,479,1,8,3,0,40,160,168]
s: 479
d: [2,494,1,8,3,0,40,172,168]
s: 494
d: [2,712,2,5,2,0,40,479,178]
s: 712
d: [2,509,1,8,2,0,40,171,173]
s: 509
d: [2,688,2,7,2,0,40,473,179]
s: 688
d: [2,386,1,9,4,1,44,382,368]
s: 386
d: [2,404,1,12,4,0,40,163,159]
d: [2,405,1,9,3,0,21,404,149]
s: 405
a: [3,655,2,5,2,648]
s: 655
r: 405
d: [2,504,1,9,3,0,40,171,168]
s: 504
d: [2,627,2,13,5,1,40,384,150]
d: [2,645,2,10,4,1,21,627,149]
d: [2,646,2,8,3,2,21,645,335]
s: 646
d: [2,506,1,9,3,0,40,171,152]
s: 506
d: [2,507,1,9,3,0,40,171,154]
s: 507
d: [2,739,2,6,2,3,21,646,655]
s: 739
d: [2,483,1,13,4,0,40,160,163]
d: [2,493,1,10,3,0,21,483,149]
s: 493
d: [2,743,2,10,3,0,40,506,419]
d: [2,747,2,7,2,0,21,743,153]
s: 747
d: [2,498,1,13,4,0,40,172,163]
d: [2,503,1,10,3,0,21,498,149]
s: 503
d: [2,513,1,10,4,0,40,150,168]
s: 513
d: [2,758,3,9,3,0,40,747,167]
d: [2,759,3,7,2,0,18,758]
s: 759
d: [2,533,1,10,4,1,44,162,249]
s: 533
d: [2,757,3,10,3,0,40,747,159]
d: [2,765,3,7,2,0,21,757,149]
s: 765
d: [2,534,1,10,4,1,44,162,254]
s: 534
d: [2,791,2,9,4,3,40,534,658]
d: [2,793,2,6,3,4,21,791,195]
d: [2,794,2,4,2,5,21,793,204]
d: [2,795,2,2,1,6,21,794,209]
d: [2,796,2,0,0,7,21,795,224]
e: 796
r: 195
a: [3,194,0,3,1,128]
s: 194
d: [2,798,2,7,3,4,21,791,204]
d: [2,799,2,5,2,5,21,798,209]
d: [2,800,2,3,1,6,21,799,224]
d: [2,802,2,7,3,2,40,194,533]
d: [2,804,2,5,2,3,21,802,204]
d: [2,805,2,3,1,4,21,804,209]
f: 805,800,643,331,137,136,133,132,131,126,127,128
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mod_2__t21_mod_2
% SZS output start Proof for mod_2__t21_mod_2
1. ! [X0] : ((v1_classes2(X0) & ~v1_xboole_0(X0)) => (m1_subset_1(2,X0) & m1_subset_1(1,X0) & m1_subset_1(k6_numbers,X0))) [input t21_mod_2]
2. ~! [X0] : ((v1_classes2(X0) & ~v1_xboole_0(X0)) => (m1_subset_1(2,X0) & m1_subset_1(1,X0) & m1_subset_1(k6_numbers,X0))) [negated conjecture 1]
23. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
33. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
36. 1 = k1_ordinal1(0) [input rqSucc__k1_ordinal1__r0_r1]
37. 2 = k1_ordinal1(1) [input rqSucc__k1_ordinal1__r1_r2]
38. v1_xboole_0(0) [input spc0_boole]
45. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
46. ! [X0] : ((v1_classes2(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_subset_1(X1,X0) => m1_subset_1(k1_ordinal1(X1),X0))) [input t20_mod_2]
50. ! [X0] : ((v1_classes2(X0) & ~v1_xboole_0(X0)) => r2_hidden(k1_xboole_0,X0)) [input t56_classes2]
52. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
55. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 45]
67. ? [X0] : ((~m1_subset_1(2,X0) | ~m1_subset_1(1,X0) | ~m1_subset_1(k6_numbers,X0)) & (v1_classes2(X0) & ~v1_xboole_0(X0))) [ennf transformation 2]
68. ? [X0] : ((~m1_subset_1(2,X0) | ~m1_subset_1(1,X0) | ~m1_subset_1(k6_numbers,X0)) & v1_classes2(X0) & ~v1_xboole_0(X0)) [flattening 67]
76. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 52]
79. ! [X0] : (r2_hidden(k1_xboole_0,X0) | (~v1_classes2(X0) | v1_xboole_0(X0))) [ennf transformation 50]
80. ! [X0] : (r2_hidden(k1_xboole_0,X0) | ~v1_classes2(X0) | v1_xboole_0(X0)) [flattening 79]
81. ! [X0] : (! [X1] : (m1_subset_1(k1_ordinal1(X1),X0) | ~m1_subset_1(X1,X0)) | (~v1_classes2(X0) | v1_xboole_0(X0))) [ennf transformation 46]
82. ! [X0] : (! [X1] : (m1_subset_1(k1_ordinal1(X1),X0) | ~m1_subset_1(X1,X0)) | ~v1_classes2(X0) | v1_xboole_0(X0)) [flattening 81]
87. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 55]
103. ? [X0] : ((~m1_subset_1(2,X0) | ~m1_subset_1(1,X0) | ~m1_subset_1(k6_numbers,X0)) & v1_classes2(X0) & ~v1_xboole_0(X0)) => ((~m1_subset_1(2,sK0) | ~m1_subset_1(1,sK0) | ~m1_subset_1(k6_numbers,sK0)) & v1_classes2(sK0) & ~v1_xboole_0(sK0)) [choice axiom]
104. (~m1_subset_1(2,sK0) | ~m1_subset_1(1,sK0) | ~m1_subset_1(k6_numbers,sK0)) & v1_classes2(sK0) & ~v1_xboole_0(sK0) [skolemisation 68,103]
126. ~v1_xboole_0(sK0) [cnf transformation 104] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
127. v1_classes2(sK0) [cnf transformation 104] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
128. ~m1_subset_1(2,sK0) | ~m1_subset_1(1,sK0) | ~m1_subset_1(k6_numbers,sK0) [cnf transformation 104] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
131. v1_xboole_0(0) [cnf transformation 38] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
132. v1_xboole_0(k1_xboole_0) [cnf transformation 23] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
133. k6_numbers = k1_xboole_0 [cnf transformation 33] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
136. 2 = k1_ordinal1(1) [cnf transformation 37] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
137. 1 = k1_ordinal1(0) [cnf transformation 36] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
158. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 76] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
161. r2_hidden(k1_xboole_0,X0) | ~v1_classes2(X0) | v1_xboole_0(X0) [cnf transformation 80] {a:0,w:7,wCS:2800,sine:2,thAx:0,allAx:1,thDist:-1}
162. m1_subset_1(k1_ordinal1(X1),X0) | ~m1_subset_1(X1,X0) | ~v1_classes2(X0) | v1_xboole_0(X0) [cnf transformation 82] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
166. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 87] {a:0,w:6,wCS:2400,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
189. 1 <=> m1_subset_1(k6_numbers,sK0) [avatar definition]
190. m1_subset_1(k6_numbers,sK0) <- (1) [avatar component clause 189] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
191. ~m1_subset_1(k6_numbers,sK0) <- (~1) [avatar component clause 189] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
193. 2 <=> m1_subset_1(1,sK0) [avatar definition]
194. m1_subset_1(1,sK0) <- (2) [avatar component clause 193] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
197. 3 <=> m1_subset_1(2,sK0) [avatar definition]
200. ~1 | ~2 | ~3 [avatar split clause 128,197,193,189]
202. 4 <=> v1_classes2(sK0) [avatar definition]
204. v1_classes2(sK0) <- (4) [avatar component clause 202] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
205. 4 [avatar split clause 127,202]
207. 5 <=> v1_xboole_0(sK0) [avatar definition]
209. ~v1_xboole_0(sK0) <- (~5) [avatar component clause 207] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
210. ~5 [avatar split clause 126,207]
222. 8 <=> v1_xboole_0(0) [avatar definition]
224. v1_xboole_0(0) <- (8) [avatar component clause 222] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
225. 8 [avatar split clause 131,222]
227. 9 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
229. v1_xboole_0(k1_xboole_0) <- (9) [avatar component clause 227] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
230. 9 [avatar split clause 132,227]
232. 10 <=> k6_numbers = k1_xboole_0 [avatar definition]
234. k6_numbers = k1_xboole_0 <- (10) [avatar component clause 232] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
235. 10 [avatar split clause 133,232]
247. 13 <=> 2 = k1_ordinal1(1) [avatar definition]
249. 2 = k1_ordinal1(1) <- (13) [avatar component clause 247] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
250. 13 [avatar split clause 136,247]
252. 14 <=> 1 = k1_ordinal1(0) [avatar definition]
254. 1 = k1_ordinal1(0) <- (14) [avatar component clause 252] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
255. 14 [avatar split clause 137,252]
331. v1_xboole_0(k6_numbers) <- (9, 10) [backward demodulation 229,234] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
333. 30 <=> v1_xboole_0(k6_numbers) [avatar definition]
335. v1_xboole_0(k6_numbers) <- (30) [avatar component clause 333] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
336. 30 | ~9 | ~10 [avatar split clause 331,232,227,333]
368. k6_numbers = X0 | ~v1_xboole_0(X0) <- (10) [forward demodulation 158,234] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
373. ~m1_subset_1(X0,sK0) | ~v1_xboole_0(X0) <- (~1, 10) [superposition 191,368] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
382. r2_hidden(k6_numbers,X0) | ~v1_classes2(X0) | v1_xboole_0(X0) <- (10) [forward demodulation 161,234] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
384. m1_subset_1(k6_numbers,X1) | v1_xboole_0(X1) | ~v1_classes2(X1) <- (10) [resolution 382,166] {a:1,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
533. ~m1_subset_1(1,X0) | m1_subset_1(2,X0) | ~v1_classes2(X0) | v1_xboole_0(X0) <- (13) [superposition 162,249] {a:1,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
534. ~m1_subset_1(0,X1) | m1_subset_1(1,X1) | ~v1_classes2(X1) | v1_xboole_0(X1) <- (14) [superposition 162,254] {a:1,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
625. v1_xboole_0(sK0) | ~v1_classes2(sK0) | ~v1_xboole_0(k6_numbers) <- (~1, 10) [resolution 384,373] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
641. ~v1_classes2(sK0) | ~v1_xboole_0(k6_numbers) <- (~1, ~5, 10) [subsumption resolution 625,209] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
642. ~v1_xboole_0(k6_numbers) <- (~1, 4, ~5, 10) [subsumption resolution 641,204] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
643. $false <- (~1, 4, ~5, 10, 30) [subsumption resolution 642,335] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
644. 1 | ~4 | 5 | ~10 | ~30 [avatar contradiction clause 643]
658. m1_subset_1(X0,sK0) | ~v1_xboole_0(X0) <- (1, 10) [superposition 190,368] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
791. m1_subset_1(1,sK0) | ~v1_classes2(sK0) | v1_xboole_0(sK0) | ~v1_xboole_0(0) <- (1, 10, 14) [resolution 534,658] {a:2,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
798. m1_subset_1(1,sK0) | v1_xboole_0(sK0) | ~v1_xboole_0(0) <- (1, 4, 10, 14) [subsumption resolution 791,204] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
799. m1_subset_1(1,sK0) | ~v1_xboole_0(0) <- (1, 4, ~5, 10, 14) [subsumption resolution 798,209] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
800. m1_subset_1(1,sK0) <- (1, 4, ~5, 8, 10, 14) [subsumption resolution 799,224] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
801. 2 | ~1 | ~4 | 5 | ~8 | ~10 | ~14 [avatar split clause 800,252,232,222,207,202,189,193]
802. m1_subset_1(2,sK0) | ~v1_classes2(sK0) | v1_xboole_0(sK0) <- (2, 13) [resolution 194,533] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
804. m1_subset_1(2,sK0) | v1_xboole_0(sK0) <- (2, 4, 13) [subsumption resolution 802,204] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
805. m1_subset_1(2,sK0) <- (2, 4, ~5, 13) [subsumption resolution 804,209] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
806. 3 | ~2 | ~4 | 5 | ~13 [avatar split clause 805,247,207,202,193,197]
813. $false [avatar sat refutation 200,205,210,225,230,235,250,255,336,644,801,806] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:279,thDist:-279}
% SZS output end Proof for mod_2__t21_mod_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 53
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 62
% Generated clauses: 462
% Active clauses: 132
% Passive clauses: 197
% Final active clauses: 122
% Final passive clauses: 51
% Main loop iterations started: 146
% 
% >>> Simplifying Inferences
% Duplicate literals: 13
% Fw subsumption resolutions: 74
% Fw demodulations: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Forward subsumptions: 105
% Backward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 235
% Unit resulting resolution: 1
% Forward superposition: 11
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 25
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 64
% SAT solver unit clauses: 27
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 5373
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% backward subsumption: 0.001 s
% resolution: 0.002 s
% other: 0.005 s ( own 0.0000 s ) 

