i: [1,36,0,2,1,1,0,0] t74_relat_1
a: [3,64,0,2,1,36]
s: 64
i: [1,37,0,3,1,1,0,0] t74_relat_1
a: [3,59,0,3,1,37]
s: 59
i: [1,38,0,9,1,1,0,0] t74_relat_1
a: [3,54,0,9,1,38]
s: 54
i: [1,45,0,3,1,0,0,1] reflexivity_r1_tarski
s: 45
i: [1,44,0,4,1,0,0,255] existence_m1_subset_1
s: 44
i: [1,46,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 46
i: [1,48,0,6,2,0,0,2] fc1_relat_1
s: 48
i: [1,49,0,6,2,0,0,1] dt_k5_relat_1
s: 49
i: [1,41,0,7,2,0,0,255] t3_subset
s: 41
i: [1,42,0,7,2,0,0,255] t3_subset
s: 42
i: [1,47,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 47
d: [2,67,1,5,1,0,40,41,44]
s: 67
i: [1,39,0,8,3,0,0,255] cc2_relat_1
s: 39
d: [2,69,1,6,2,0,44,48,47]
s: 69
i: [1,40,0,8,2,0,0,1] t28_xboole_1
s: 40
i: [1,50,0,11,2,0,0,1] rd5_relat_1
s: 50
d: [2,70,1,6,2,0,44,48,47]
s: 70
i: [1,43,0,13,2,0,0,1] t71_relat_1
s: 43
d: [2,87,1,11,2,1,44,54,43]
d: [2,95,1,9,1,2,21,87,64]
d: [2,96,1,9,1,2,22,95,47]
a: [3,100,1,9,1,96]
s: 100
d: [2,103,2,10,2,1,44,100,40]
d: [2,104,2,3,1,1,19,103]
d: [2,105,2,0,0,2,21,104,59]
e: 105
f: 105,96,36,37,38
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t74_relat_1
% SZS output start Proof for relat_1__t74_relat_1
1. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => k5_relat_1(k5_relat_1(X2,X1),X0) = k5_relat_1(X2,X0))) [input t74_relat_1]
2. ~! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => k5_relat_1(k5_relat_1(X2,X1),X0) = k5_relat_1(X2,X0))) [negated conjecture 1]
4. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
14. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
16. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1))) [input t71_relat_1]
17. ~! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(X0,X1) => k5_relat_1(k5_relat_1(X2,X1),X0) = k5_relat_1(X2,X0))) [flattening 2]
18. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 14]
20. ! [X0,X1,X2] : (v1_relat_1(X2) => k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1))) [flattening 16]
23. ? [X0,X1,X2] : ((k5_relat_1(k5_relat_1(X2,X1),X0) != k5_relat_1(X2,X0) & r1_tarski(X0,X1)) & v1_relat_1(X2)) [ennf transformation 17]
24. ? [X0,X1,X2] : (k5_relat_1(k5_relat_1(X2,X1),X0) != k5_relat_1(X2,X0) & r1_tarski(X0,X1) & v1_relat_1(X2)) [flattening 23]
26. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 18]
27. ! [X0,X1,X2] : (k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1)) | ~v1_relat_1(X2)) [ennf transformation 20]
31. ? [X0,X1,X2] : (k5_relat_1(k5_relat_1(X2,X1),X0) != k5_relat_1(X2,X0) & r1_tarski(X0,X1) & v1_relat_1(X2)) => (k5_relat_1(k5_relat_1(sK2,sK1),sK0) != k5_relat_1(sK2,sK0) & r1_tarski(sK0,sK1) & v1_relat_1(sK2)) [choice axiom]
32. k5_relat_1(k5_relat_1(sK2,sK1),sK0) != k5_relat_1(sK2,sK0) & r1_tarski(sK0,sK1) & v1_relat_1(sK2) [skolemisation 24,31]
36. v1_relat_1(sK2) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. r1_tarski(sK0,sK1) [cnf transformation 32] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. k5_relat_1(k5_relat_1(sK2,sK1),sK0) != k5_relat_1(sK2,sK0) [cnf transformation 32] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1)) | ~v1_relat_1(X2) [cnf transformation 27] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
52. 1 <=> k5_relat_1(k5_relat_1(sK2,sK1),sK0) = k5_relat_1(sK2,sK0) [avatar definition]
54. k5_relat_1(k5_relat_1(sK2,sK1),sK0) != k5_relat_1(sK2,sK0) <- (~1) [avatar component clause 52] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
55. ~1 [avatar split clause 38,52]
57. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
59. r1_tarski(sK0,sK1) <- (2) [avatar component clause 57] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 2 [avatar split clause 37,57]
62. 3 <=> v1_relat_1(sK2) [avatar definition]
64. v1_relat_1(sK2) <- (3) [avatar component clause 62] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. 3 [avatar split clause 36,62]
87. k5_relat_1(sK2,sK0) != k5_relat_1(sK2,k3_xboole_0(sK1,sK0)) | ~v1_relat_1(sK2) <- (~1) [superposition 54,43] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
95. k5_relat_1(sK2,sK0) != k5_relat_1(sK2,k3_xboole_0(sK1,sK0)) <- (~1, 3) [subsumption resolution 87,64] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
96. k5_relat_1(sK2,sK0) != k5_relat_1(sK2,k3_xboole_0(sK0,sK1)) <- (~1, 3) [forward demodulation 95,47] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
98. 4 <=> k5_relat_1(sK2,sK0) = k5_relat_1(sK2,k3_xboole_0(sK0,sK1)) [avatar definition]
100. k5_relat_1(sK2,sK0) != k5_relat_1(sK2,k3_xboole_0(sK0,sK1)) <- (~4) [avatar component clause 98] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. ~4 | 1 | ~3 [avatar split clause 96,62,52,98]
103. k5_relat_1(sK2,sK0) != k5_relat_1(sK2,sK0) | ~r1_tarski(sK0,sK1) <- (~4) [superposition 100,40] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
104. ~r1_tarski(sK0,sK1) <- (~4) [trivial inequality removal 103] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
105. $false <- (2, ~4) [subsumption resolution 104,59] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
106. ~2 | 4 [avatar contradiction clause 105]
107. $false [avatar sat refutation 55,60,65,101,106] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:19,thDist:-19}
% SZS output end Proof for relat_1__t74_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 54
% Active clauses: 18
% Passive clauses: 29
% Final active clauses: 18
% Final passive clauses: 10
% Main loop iterations started: 19
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Trivial inequalities: 1
% Fw subsumption resolutions: 5
% Fw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 5
% Fw demodulations to eq. taut.: 2
% 
% >>> Generating Inferences
% Binary resolution: 4
% Forward superposition: 9
% Backward superposition: 11
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 7
% SAT solver unit clauses: 3
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

