i: [1,36,0,26,9,1,0,0] t36_gate_1
a: [3,58,0,2,1,36]
s: 58
i: [1,37,0,2,1,0,0,255] fc1_xboole_0
a: [3,99,0,2,1,37]
s: 99
i: [1,49,0,2,1,0,0,1] rc2_xboole_0
a: [3,104,0,2,1,49]
s: 104
i: [1,50,0,2,1,0,0,1] rc1_xboole_0
a: [3,109,0,2,1,50]
s: 109
a: [3,53,0,10,1,36]
s: 53
i: [1,38,0,5,2,0,0,255] fc2_gate_1
s: 38
i: [1,39,0,5,2,0,0,255] fc1_gate_1
s: 39
i: [1,41,0,13,2,0,0,255] d30_gate_1
s: 41
d: [2,112,1,4,2,1,44,53,41]
d: [2,113,1,2,1,2,21,112,99]
d: [2,114,1,0,0,3,21,113,58]
e: 114
r: 58
a: [3,57,0,2,1,36]
s: 57
a: [3,70,0,2,1,36]
s: 70
i: [1,42,0,13,2,0,0,255] d30_gate_1
s: 42
a: [3,61,0,2,1,36]
s: 61
i: [1,43,0,13,2,0,0,255] d30_gate_1
s: 43
a: [3,65,0,2,1,36]
s: 65
i: [1,44,0,13,2,0,0,255] d30_gate_1
s: 44
d: [2,123,1,4,2,1,44,53,44]
d: [2,124,1,2,1,2,21,123,99]
d: [2,125,1,0,0,3,21,124,70]
e: 125
r: 70
a: [3,69,0,2,1,36]
s: 69
a: [3,74,0,2,1,36]
s: 74
i: [1,45,0,13,2,0,0,255] d30_gate_1
s: 45
d: [2,128,1,4,2,1,44,53,45]
d: [2,129,1,2,1,2,21,128,99]
d: [2,130,1,0,0,3,21,129,74]
e: 130
r: 74
a: [3,73,0,2,1,36]
s: 73
a: [3,82,0,2,1,36]
s: 82
i: [1,46,0,13,2,0,0,255] d30_gate_1
s: 46
a: [3,77,0,2,1,36]
s: 77
i: [1,47,0,13,2,0,0,255] d30_gate_1
s: 47
d: [2,136,1,4,2,1,44,53,47]
d: [2,137,1,2,1,2,21,136,99]
d: [2,138,1,0,0,3,21,137,82]
e: 138
r: 82
a: [3,81,0,2,1,36]
s: 81
a: [3,86,0,2,1,36]
s: 86
i: [1,48,0,13,2,0,0,255] d30_gate_1
s: 48
d: [2,141,1,4,2,1,44,53,48]
d: [2,142,1,2,1,2,21,141,99]
d: [2,143,1,0,0,3,21,142,86]
e: 143
r: 53
r: 86
a: [3,85,0,2,1,36]
s: 85
a: [3,54,0,10,1,36]
s: 54
i: [1,40,0,28,9,0,0,255] d30_gate_1
s: 40
d: [2,169,1,19,9,1,44,54,40]
d: [2,170,1,17,8,2,21,169,85]
d: [2,171,1,15,7,3,21,170,81]
d: [2,172,1,13,6,4,21,171,77]
d: [2,173,1,11,5,5,21,172,73]
d: [2,174,1,9,4,6,21,173,69]
d: [2,175,1,7,3,7,21,174,65]
d: [2,176,1,5,2,8,21,175,61]
d: [2,177,1,3,1,9,21,176,57]
a: [3,181,1,3,1,177]
s: 181
d: [2,183,2,0,0,2,52,99,181,39]
e: 183
d: [2,184,2,2,1,1,40,181,39]
d: [2,186,2,0,0,2,21,184,99]
e: 186
i: [1,35,0,12,2,1,0,0] t36_gate_1
i: [1,34,0,12,2,1,0,0] t36_gate_1
i: [1,33,0,12,2,1,0,0] t36_gate_1
i: [1,32,0,12,2,1,0,0] t36_gate_1
i: [1,31,0,12,2,1,0,0] t36_gate_1
i: [1,30,0,12,2,1,0,0] t36_gate_1
i: [1,29,0,12,2,1,0,0] t36_gate_1
i: [1,28,0,12,2,1,0,0] t36_gate_1
d: [2,117,1,4,2,1,44,53,42]
d: [2,118,1,2,1,2,21,117,99]
d: [2,120,1,4,2,1,44,53,43]
d: [2,121,1,2,1,2,21,120,99]
d: [2,133,1,4,2,1,44,53,46]
d: [2,134,1,2,1,2,21,133,99]
f: 186,177,143,137,134,129,124,121,118,113,37,28,29,30,31,32,33,34,35,36
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t36_gate_1
% SZS output start Proof for gate_1__t36_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (~v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)) <=> (~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input t36_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (~v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)) <=> (~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ((~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0) & ~(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input d30_gate_1]
14. ~! [X0,X1,X2,X3,X4,X5,X6,X7] : (~v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)) <=> (~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 2]
15. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0) & ~(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 13]
16. ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)) <~> (~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | (~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & (k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X7) | v1_xboole_0(X6) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 15]
20. ? [X0,X1,X2,X3,X4,X5,X6,X7] : (((v1_xboole_0(X7) | v1_xboole_0(X6) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) | v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7))) & ((~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)))) [nnf transformation 16]
21. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((v1_xboole_0(X7) | v1_xboole_0(X6) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7))) & ((~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)))) [flattening 20]
22. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((v1_xboole_0(X7) | v1_xboole_0(X6) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7))) & ((~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)))) => ((v1_xboole_0(sK7) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7))) & ((~v1_xboole_0(sK7) & ~v1_xboole_0(sK6) & ~v1_xboole_0(sK5) & ~v1_xboole_0(sK4) & ~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)))) [choice axiom]
23. (v1_xboole_0(sK7) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7))) & ((~v1_xboole_0(sK7) & ~v1_xboole_0(sK6) & ~v1_xboole_0(sK5) & ~v1_xboole_0(sK4) & ~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7))) [skolemisation 21,22]
28. ~v1_xboole_0(sK0) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 23] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~v1_xboole_0(sK1) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 23] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. ~v1_xboole_0(sK2) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 23] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~v1_xboole_0(sK3) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 23] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. ~v1_xboole_0(sK4) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 23] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. ~v1_xboole_0(sK5) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 23] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. ~v1_xboole_0(sK6) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 23] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
35. ~v1_xboole_0(sK7) | ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 23] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. v1_xboole_0(sK7) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 23] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
39. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X7) | v1_xboole_0(X6) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:28,wCS:11200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
41. k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | ~v1_xboole_0(X1) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | ~v1_xboole_0(X2) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | ~v1_xboole_0(X3) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
45. k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | ~v1_xboole_0(X4) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
46. k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | ~v1_xboole_0(X5) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
47. k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | ~v1_xboole_0(X6) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
48. k30_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | ~v1_xboole_0(X7) [cnf transformation 19] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
52. 1 <=> v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [avatar definition]
53. ~v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) <- (~1) [avatar component clause 52] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
54. v1_xboole_0(k30_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) <- (1) [avatar component clause 52] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. 2 <=> v1_xboole_0(sK0) [avatar definition]
57. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
60. 3 <=> v1_xboole_0(sK1) [avatar definition]
61. ~v1_xboole_0(sK1) <- (~3) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
64. 4 <=> v1_xboole_0(sK2) [avatar definition]
65. ~v1_xboole_0(sK2) <- (~4) [avatar component clause 64] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. 5 <=> v1_xboole_0(sK3) [avatar definition]
69. ~v1_xboole_0(sK3) <- (~5) [avatar component clause 68] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. 6 <=> v1_xboole_0(sK4) [avatar definition]
73. ~v1_xboole_0(sK4) <- (~6) [avatar component clause 72] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. 7 <=> v1_xboole_0(sK5) [avatar definition]
77. ~v1_xboole_0(sK5) <- (~7) [avatar component clause 76] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. 8 <=> v1_xboole_0(sK6) [avatar definition]
81. ~v1_xboole_0(sK6) <- (~8) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. 9 <=> v1_xboole_0(sK7) [avatar definition]
85. ~v1_xboole_0(sK7) <- (~9) [avatar component clause 84] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
86. v1_xboole_0(sK7) <- (9) [avatar component clause 84] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. 1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | 9 [avatar split clause 36,84,80,76,72,68,64,60,56,52]
88. ~1 | ~9 [avatar split clause 35,84,52]
89. ~1 | ~8 [avatar split clause 34,80,52]
90. ~1 | ~7 [avatar split clause 33,76,52]
91. ~1 | ~6 [avatar split clause 32,72,52]
92. ~1 | ~5 [avatar split clause 31,68,52]
93. ~1 | ~4 [avatar split clause 30,64,52]
94. ~1 | ~3 [avatar split clause 29,60,52]
95. ~1 | ~2 [avatar split clause 28,56,52]
97. 10 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
99. v1_xboole_0(k1_xboole_0) <- (10) [avatar component clause 97] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
100. 10 [avatar split clause 37,97]
112. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK0) <- (~1) [superposition 53,41] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. ~v1_xboole_0(sK0) <- (~1, 10) [subsumption resolution 112,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. ~2 | 1 | ~10 [avatar split clause 113,97,52,56]
117. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK1) <- (~1) [superposition 53,42] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
118. ~v1_xboole_0(sK1) <- (~1, 10) [subsumption resolution 117,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. ~3 | 1 | ~10 [avatar split clause 118,97,52,60]
120. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK2) <- (~1) [superposition 53,43] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
121. ~v1_xboole_0(sK2) <- (~1, 10) [subsumption resolution 120,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
122. ~4 | 1 | ~10 [avatar split clause 121,97,52,64]
123. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK3) <- (~1) [superposition 53,44] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. ~v1_xboole_0(sK3) <- (~1, 10) [subsumption resolution 123,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
127. ~5 | 1 | ~10 [avatar split clause 124,97,52,68]
128. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK4) <- (~1) [superposition 53,45] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. ~v1_xboole_0(sK4) <- (~1, 10) [subsumption resolution 128,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
132. ~6 | 1 | ~10 [avatar split clause 129,97,52,72]
133. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK5) <- (~1) [superposition 53,46] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. ~v1_xboole_0(sK5) <- (~1, 10) [subsumption resolution 133,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
135. ~7 | 1 | ~10 [avatar split clause 134,97,52,76]
136. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK6) <- (~1) [superposition 53,47] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
137. ~v1_xboole_0(sK6) <- (~1, 10) [subsumption resolution 136,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
140. ~8 | 1 | ~10 [avatar split clause 137,97,52,80]
141. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK7) <- (~1) [superposition 53,48] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
142. ~v1_xboole_0(sK7) <- (~1, 10) [subsumption resolution 141,99] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
143. $false <- (~1, 9, 10) [subsumption resolution 142,86] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. 1 | ~9 | ~10 [avatar contradiction clause 143]
169. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK7) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1) [superposition 54,40] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
170. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~9) [subsumption resolution 169,85] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
171. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~8, ~9) [subsumption resolution 170,81] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
172. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~7, ~8, ~9) [subsumption resolution 171,77] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
173. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~6, ~7, ~8, ~9) [subsumption resolution 172,73] {a:1,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
174. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~5, ~6, ~7, ~8, ~9) [subsumption resolution 173,69] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
175. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~4, ~5, ~6, ~7, ~8, ~9) [subsumption resolution 174,65] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
176. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK0) <- (1, ~3, ~4, ~5, ~6, ~7, ~8, ~9) [subsumption resolution 175,61] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
177. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (1, ~2, ~3, ~4, ~5, ~6, ~7, ~8, ~9) [subsumption resolution 176,57] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
179. 13 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
181. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (13) [avatar component clause 179] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
182. 13 | ~1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | 9 [avatar split clause 177,84,80,76,72,68,64,60,56,52,179]
184. ~v1_xboole_0(k1_xboole_0) <- (13) [resolution 181,39] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
186. $false <- (10, 13) [subsumption resolution 184,99] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
187. ~10 | ~13 [avatar contradiction clause 186]
188. $false [avatar sat refutation 87,88,89,90,91,92,93,94,95,100,116,119,122,127,132,135,140,144,182,187] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:130,thDist:-130}
% SZS output end Proof for gate_1__t36_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 10
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 104
% Active clauses: 30
% Passive clauses: 30
% Final active clauses: 24
% Main loop iterations started: 30
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 22
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 16
% Backward superposition: 17
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 25
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 29
% SAT solver unit clauses: 3
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward superposition index maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

