i: [1,83,0,3,1,1,0,0] t34_subset_1
a: [3,140,0,3,1,83]
s: 140
i: [1,82,0,3,1,1,0,0] t34_subset_1
a: [3,145,0,3,1,82]
s: 145
i: [1,81,0,3,1,1,0,0] t34_subset_1
a: [3,150,0,3,1,81]
s: 150
i: [1,84,0,6,1,1,0,0] t34_subset_1
a: [3,135,0,6,1,84]
s: 135
i: [1,85,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,155,0,2,1,85]
s: 155
i: [1,86,0,2,1,0,0,1] fc1_xboole_0
a: [3,160,0,2,1,86]
s: 160
i: [1,120,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,170,0,2,1,120]
s: 170
i: [1,121,0,2,1,0,0,2] rc2_xboole_0
a: [3,175,0,2,1,121]
s: 175
i: [1,122,0,2,1,0,0,2] rc1_xboole_0
a: [3,180,0,2,1,122]
s: 180
i: [1,124,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,185,0,2,1,124]
s: 185
i: [1,123,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,190,0,2,1,123]
s: 190
i: [1,88,0,3,1,0,0,1] fc1_subset_1
s: 88
i: [1,115,0,3,1,0,0,1] rc2_subset_1
s: 115
i: [1,116,0,3,1,0,0,255] reflexivity_r1_tarski
s: 116
i: [1,87,0,3,1,0,0,255] d2_xboole_0
a: [3,165,0,3,1,87]
s: 165
r: 155
i: [1,89,0,4,2,0,0,255] cc2_zfmisc_1
s: 89
d: [2,193,1,2,1,1,40,89,160]
a: [3,199,1,2,1,193]
s: 199
i: [1,92,0,4,2,0,0,255] cc1_zfmisc_1
s: 92
d: [2,195,1,2,1,1,40,89,180]
a: [3,204,1,2,1,195]
s: 204
i: [1,113,0,4,1,0,0,1] existence_m1_subset_1
s: 113
i: [1,117,0,4,1,0,0,1] fc3_xboole_0
s: 117
d: [2,194,1,3,1,0,40,89,115]
s: 194
i: [1,91,0,5,2,0,0,1] rc1_subset_1
s: 91
i: [1,93,0,5,2,0,0,1] t6_boole
s: 93
i: [1,106,0,5,2,0,0,2] t7_boole
s: 106
i: [1,114,0,5,1,0,0,1] rc2_subset_1
s: 114
i: [1,128,0,5,1,0,0,1] d2_tarski
s: 128
i: [1,130,0,5,1,0,0,1] d2_tarski
s: 130
i: [1,119,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 119
i: [1,90,0,7,2,0,0,1] rc1_subset_1
s: 90
i: [1,96,0,7,3,0,0,1] d1_subset_1
s: 96
d: [2,213,1,4,2,1,40,96,145]
a: [3,221,1,2,1,213]
s: 221
i: [1,97,0,7,3,0,0,1] d1_subset_1
s: 97
d: [2,211,1,5,1,0,40,119,128]
s: 211
i: [1,105,0,7,3,0,0,1] t8_boole
s: 105
i: [1,118,0,7,1,0,0,1] commutativity_k2_tarski
s: 118
d: [2,212,1,5,1,0,40,119,130]
s: 212
i: [1,125,0,7,2,0,0,255] d1_zfmisc_1
s: 125
d: [2,215,1,5,2,0,40,96,113]
s: 215
i: [1,126,0,7,2,0,0,255] d1_zfmisc_1
s: 126
i: [1,94,0,8,3,0,0,1] d1_subset_1
s: 94
d: [2,255,1,5,2,1,40,94,145]
d: [2,259,1,3,1,2,21,255,221]
a: [3,263,1,3,1,259]
s: 263
i: [1,95,0,8,3,0,0,1] d1_subset_1
s: 95
d: [2,256,1,5,2,1,40,94,150]
d: [2,265,1,3,1,2,21,256,221]
a: [3,269,1,3,1,265]
s: 269
d: [2,278,0,6,2,0,21,95,106]
s: 278
i: [1,99,0,8,2,0,0,255] d3_tarski
s: 99
d: [2,271,2,3,1,1,40,263,119]
a: [3,276,2,3,1,271]
s: 276
i: [1,100,0,8,2,0,0,255] d3_tarski
s: 100
d: [2,279,2,3,1,1,40,269,119]
a: [3,284,2,3,1,279]
s: 284
i: [1,98,0,9,3,0,0,255] d3_tarski
s: 98
i: [1,131,0,11,3,0,0,1] d2_tarski
s: 131
d: [2,253,1,8,2,0,40,94,114]
d: [2,257,1,5,1,0,21,253,88]
s: 257
i: [1,103,0,14,3,0,0,255] d1_zfmisc_1
s: 103
d: [2,304,2,4,1,0,40,257,126]
s: 304
i: [1,104,0,14,3,0,0,255] d1_zfmisc_1
s: 104
i: [1,111,0,17,3,0,0,1] d2_tarski
s: 111
d: [2,286,1,5,1,0,40,278,128]
s: 286
i: [1,112,0,17,3,0,0,1] d2_tarski
s: 112
d: [2,287,1,5,1,0,40,278,130]
s: 287
i: [1,110,0,23,4,0,0,1] d2_tarski
s: 110
d: [2,293,1,5,2,0,40,99,106]
s: 293
d: [2,235,2,5,1,0,44,211,118]
s: 235
d: [2,251,1,6,2,0,40,94,113]
s: 251
d: [2,239,2,5,1,0,44,211,118]
s: 239
d: [2,206,1,7,3,0,44,93,93]
s: 206
d: [2,245,1,7,2,0,40,125,119]
s: 245
d: [2,249,2,5,2,0,40,215,89]
s: 249
d: [2,254,1,10,3,0,40,94,90]
d: [2,258,1,7,2,0,21,254,88]
s: 258
d: [2,306,2,5,1,0,40,257,119]
s: 306
d: [2,288,1,7,2,0,40,278,125]
s: 288
d: [2,291,1,8,2,0,40,99,278]
s: 291
d: [2,341,2,5,1,1,40,288,135]
a: [3,347,2,5,1,341]
s: 347
d: [2,292,1,8,2,0,40,99,119]
s: 292
d: [2,314,2,5,1,0,44,286,118]
s: 314
d: [2,294,1,10,2,0,40,99,126]
s: 294
d: [2,296,1,10,2,0,40,100,125]
s: 296
d: [2,315,2,5,1,0,44,286,118]
s: 315
d: [2,308,1,17,4,0,40,103,98]
s: 308
d: [2,331,2,8,2,0,40,251,126]
d: [2,333,2,5,1,0,21,331,88]
s: 333
d: [2,301,1,19,3,0,40,131,99]
s: 301
d: [2,322,1,19,4,0,40,110,106]
s: 322
d: [2,247,2,6,1,0,40,125,211]
s: 247
d: [2,320,1,23,4,0,40,110,278]
s: 320
d: [2,248,2,6,1,0,40,125,212]
s: 248
d: [2,321,1,23,4,0,40,110,119]
s: 321
d: [2,324,1,27,4,0,40,110,126]
s: 324
d: [2,328,2,6,2,0,40,251,119]
s: 328
d: [2,323,1,39,5,0,40,110,131]
s: 323
d: [2,334,2,6,1,0,40,245,128]
s: 334
d: [2,335,2,6,1,0,40,245,130]
s: 335
d: [2,337,2,6,2,0,40,258,126]
s: 337
d: [2,381,3,6,1,0,44,247,118]
s: 381
d: [2,339,2,7,2,0,40,258,119]
s: 339
d: [2,382,3,6,1,0,44,247,118]
s: 382
d: [2,336,2,8,2,0,40,245,125]
s: 336
d: [2,354,2,9,2,0,40,296,293]
s: 354
d: [2,413,3,5,1,0,40,336,116]
s: 413
d: [2,330,2,16,3,0,40,251,131]
d: [2,332,2,12,2,0,21,330,117]
s: 332
d: [2,309,3,7,2,0,40,304,98]
s: 309
d: [2,326,2,13,3,0,40,293,104]
s: 326
d: [2,352,2,13,3,0,40,294,98]
s: 352
d: [2,359,3,8,2,0,40,333,98]
s: 359
d: [2,367,2,15,3,0,46,301]
s: 367
d: [2,412,3,9,3,0,40,337,98]
s: 412
d: [2,364,2,20,4,0,44,100,301]
d: [2,370,2,15,3,0,18,364]
s: 370
d: [2,376,2,16,4,0,46,322]
s: 376
d: [2,440,4,10,3,0,40,412,251]
d: [2,441,4,7,2,0,21,440,91]
s: 441
d: [2,366,2,23,4,0,44,296,301]
d: [2,368,2,17,3,0,18,366]
s: 368
d: [2,464,5,9,2,0,40,441,126]
d: [2,471,5,6,1,0,21,464,88]
s: 471
d: [2,358,2,19,4,0,40,308,251]
s: 358
d: [2,356,2,21,4,0,40,308,99]
s: 356
d: [2,460,5,7,2,0,40,441,278]
s: 460
d: [2,390,2,30,5,0,40,324,98]
s: 390
d: [2,461,5,7,2,0,40,441,119]
s: 461
d: [2,401,2,34,5,0,46,323]
s: 401
d: [2,402,2,34,5,0,46,323]
s: 402
d: [2,465,5,11,3,0,40,441,412]
d: [2,472,5,8,2,0,21,465,91]
s: 472
d: [2,403,2,34,5,0,46,323]
s: 403
d: [2,529,6,10,2,0,40,472,126]
d: [2,537,6,7,1,0,21,529,88]
s: 537
d: [2,399,2,46,7,0,44,112,323]
d: [2,405,2,39,6,0,18,399]
d: [2,410,2,34,5,0,21,405,128]
s: 410
d: [2,400,2,46,7,0,44,111,323]
d: [2,404,2,39,6,0,18,400]
d: [2,411,2,34,5,0,21,404,128]
s: 411
d: [2,525,6,8,2,0,40,472,278]
s: 525
d: [2,357,2,40,6,0,40,308,110]
s: 357
d: [2,526,6,8,2,0,40,472,119]
s: 526
d: [2,425,3,9,2,0,46,332]
s: 425
d: [2,429,3,16,3,0,40,326,125]
d: [2,430,3,10,2,0,21,429,293]
s: 430
d: [2,556,4,6,1,0,47,425]
s: 556
d: [2,447,3,16,4,0,44,100,370]
d: [2,455,3,11,3,0,18,447]
s: 455
d: [2,562,4,0,0,3,52,263,269,347,455]
e: 562
d: [2,563,4,6,2,1,40,455,347]
d: [2,570,4,3,1,2,21,563,269]
d: [2,571,4,0,0,3,21,570,263]
e: 571
r: 221
r: 269
r: 263
a: [3,220,1,2,1,213]
s: 220
d: [2,579,2,0,0,2,52,140,220,93]
e: 579
d: [2,580,2,0,0,3,52,160,140,220,105]
e: 580
d: [2,581,2,0,0,3,52,160,140,220,105]
e: 581
f: 581,255,563,256,341,86,81,82,83,84
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t34_subset_1
% SZS output start Proof for subset_1__t34_subset_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_subset_1(X2,X0) => (k1_xboole_0 != X0 => m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(X0))))) [input t34_subset_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_subset_1(X2,X0) => (k1_xboole_0 != X0 => m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(X0))))) [negated conjecture 1]
7. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
8. ! [X0] : ! [X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [input d1_zfmisc_1]
9. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
11. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
19. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
29. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
30. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [input t8_boole]
31. ~! [X0,X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_subset_1(X2,X0) => (k1_xboole_0 != X0 => m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(X0))))) [flattening 2]
32. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 7]
33. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 11]
34. ! [X0,X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [flattening 8]
35. ! [X0,X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [flattening 30]
36. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 29]
37. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 9]
39. ? [X0,X1] : (? [X2] : ((~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(X0)) & k1_xboole_0 != X0) & m1_subset_1(X2,X0)) & m1_subset_1(X1,X0)) [ennf transformation 31]
40. ? [X0,X1] : (? [X2] : (~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(X0)) & k1_xboole_0 != X0 & m1_subset_1(X2,X0)) & m1_subset_1(X1,X0)) [flattening 39]
45. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 32]
46. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 33]
47. ! [X0,X1] : (~v1_xboole_0(X1) | X0 = X1 | ~v1_xboole_0(X0)) [ennf transformation 35]
48. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 36]
50. ? [X0,X1] : (? [X2] : (~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(X0)) & k1_xboole_0 != X0 & m1_subset_1(X2,X0)) & m1_subset_1(X1,X0)) => (? [X2] : (~m1_subset_1(k2_tarski(sK1,X2),k1_zfmisc_1(sK0)) & k1_xboole_0 != sK0 & m1_subset_1(X2,sK0)) & m1_subset_1(sK1,sK0)) [choice axiom]
51. ? [X2] : (~m1_subset_1(k2_tarski(sK1,X2),k1_zfmisc_1(sK0)) & k1_xboole_0 != sK0 & m1_subset_1(X2,sK0)) => (~m1_subset_1(k2_tarski(sK1,sK2),k1_zfmisc_1(sK0)) & k1_xboole_0 != sK0 & m1_subset_1(sK2,sK0)) [choice axiom]
52. (~m1_subset_1(k2_tarski(sK1,sK2),k1_zfmisc_1(sK0)) & k1_xboole_0 != sK0 & m1_subset_1(sK2,sK0)) & m1_subset_1(sK1,sK0) [skolemisation 40,51,50]
55. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 45]
56. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 46]
57. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 56]
58. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK4(X0,X1),X0))) [choice axiom]
59. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK4(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 57,58]
60. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ~r1_tarski(X2,X0)) & (r1_tarski(X2,X0) | ~r2_hidden(X2,X1))) | k1_zfmisc_1(X0) != X1)) [nnf transformation 34]
61. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [rectify 60]
62. ! [X1,X0] : (? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1))) => ((~r1_tarski(sK5(X0,X1),X0) | ~r2_hidden(sK5(X0,X1),X1)) & (r1_tarski(sK5(X0,X1),X0) | r2_hidden(sK5(X0,X1),X1)))) [choice axiom]
63. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ((~r1_tarski(sK5(X0,X1),X0) | ~r2_hidden(sK5(X0,X1),X1)) & (r1_tarski(sK5(X0,X1),X0) | r2_hidden(sK5(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [skolemisation 61,62]
64. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 37]
65. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 64]
66. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 65]
67. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK6(X0,X1,X2) != X1 & sK6(X0,X1,X2) != X0) | ~r2_hidden(sK6(X0,X1,X2),X2)) & (sK6(X0,X1,X2) = X1 | sK6(X0,X1,X2) = X0 | r2_hidden(sK6(X0,X1,X2),X2)))) [choice axiom]
68. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK6(X0,X1,X2) != X1 & sK6(X0,X1,X2) != X0) | ~r2_hidden(sK6(X0,X1,X2),X2)) & (sK6(X0,X1,X2) = X1 | sK6(X0,X1,X2) = X0 | r2_hidden(sK6(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 66,67]
81. m1_subset_1(sK1,sK0) [cnf transformation 52] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
82. m1_subset_1(sK2,sK0) [cnf transformation 52] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
83. k1_xboole_0 != sK0 [cnf transformation 52] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
84. ~m1_subset_1(k2_tarski(sK1,sK2),k1_zfmisc_1(sK0)) [cnf transformation 52] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
86. v1_xboole_0(k1_xboole_0) [cnf transformation 19] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
94. ~m1_subset_1(X1,X0) | r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 55] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
95. m1_subset_1(X1,X0) | ~r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 55] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
99. r2_hidden(sK4(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 59] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
100. ~r2_hidden(sK4(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 59] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
102. r2_hidden(X3,X1) | ~r1_tarski(X3,X0) | k1_zfmisc_1(X0) != X1 [cnf transformation 63] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
105. X0 = X1 | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 47] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
106. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 48] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
107. X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2) | k2_tarski(X0,X1) != X2 [cnf transformation 68] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
125. r2_hidden(X3,k1_zfmisc_1(X0)) | ~r1_tarski(X3,X0) [equality resolution 102] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
131. ~r2_hidden(X4,k2_tarski(X0,X1)) | X0 = X4 | X1 = X4 [equality resolution 107] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
133. 1 <=> m1_subset_1(k2_tarski(sK1,sK2),k1_zfmisc_1(sK0)) [avatar definition]
135. ~m1_subset_1(k2_tarski(sK1,sK2),k1_zfmisc_1(sK0)) <- (~1) [avatar component clause 133] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
136. ~1 [avatar split clause 84,133]
138. 2 <=> k1_xboole_0 = sK0 [avatar definition]
140. k1_xboole_0 != sK0 <- (~2) [avatar component clause 138] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
141. ~2 [avatar split clause 83,138]
143. 3 <=> m1_subset_1(sK2,sK0) [avatar definition]
145. m1_subset_1(sK2,sK0) <- (3) [avatar component clause 143] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
146. 3 [avatar split clause 82,143]
148. 4 <=> m1_subset_1(sK1,sK0) [avatar definition]
150. m1_subset_1(sK1,sK0) <- (4) [avatar component clause 148] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
151. 4 [avatar split clause 81,148]
158. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
160. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 158] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
161. 6 [avatar split clause 86,158]
219. 15 <=> v1_xboole_0(sK0) [avatar definition]
220. v1_xboole_0(sK0) <- (15) [avatar component clause 219] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
255. r2_hidden(sK2,sK0) | v1_xboole_0(sK0) <- (3) [resolution 94,145] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
256. r2_hidden(sK1,sK0) | v1_xboole_0(sK0) <- (4) [resolution 94,150] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
261. 18 <=> r2_hidden(sK2,sK0) [avatar definition]
267. 19 <=> r2_hidden(sK1,sK0) [avatar definition]
278. ~r2_hidden(X1,X0) | m1_subset_1(X1,X0) [subsumption resolution 95,106] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
288. m1_subset_1(X4,k1_zfmisc_1(X5)) | ~r1_tarski(X4,X5) [resolution 278,125] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
301. sK4(k2_tarski(X4,X5),X6) = X5 | sK4(k2_tarski(X4,X5),X6) = X4 | r1_tarski(k2_tarski(X4,X5),X6) [resolution 131,99] {a:1,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
341. ~r1_tarski(k2_tarski(sK1,sK2),sK0) <- (~1) [resolution 288,135] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
345. 22 <=> r1_tarski(k2_tarski(sK1,sK2),sK0) [avatar definition]
347. ~r1_tarski(k2_tarski(sK1,sK2),sK0) <- (~22) [avatar component clause 345] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
348. ~22 | 1 [avatar split clause 341,133,345]
364. ~r2_hidden(X7,X8) | r1_tarski(k2_tarski(X6,X7),X8) | sK4(k2_tarski(X6,X7),X8) = X6 | r1_tarski(k2_tarski(X6,X7),X8) [superposition 100,301] {a:2,w:20,wCS:8000,sine:1,thAx:0,allAx:6,thDist:-6}
370. sK4(k2_tarski(X6,X7),X8) = X6 | r1_tarski(k2_tarski(X6,X7),X8) | ~r2_hidden(X7,X8) [duplicate literal removal 364] {a:2,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
447. ~r2_hidden(X6,X8) | r1_tarski(k2_tarski(X6,X7),X8) | r1_tarski(k2_tarski(X6,X7),X8) | ~r2_hidden(X7,X8) [superposition 100,370] {a:3,w:16,wCS:6400,sine:1,thAx:0,allAx:8,thDist:-8}
455. r1_tarski(k2_tarski(X6,X7),X8) | ~r2_hidden(X6,X8) | ~r2_hidden(X7,X8) [duplicate literal removal 447] {a:3,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:8,thDist:-8}
563. ~r2_hidden(sK1,sK0) | ~r2_hidden(sK2,sK0) <- (~22) [resolution 455,347] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
574. 15 | 19 | ~4 [avatar split clause 256,148,267,219]
577. ~18 | ~19 | 22 [avatar split clause 563,345,267,261]
578. 15 | 18 | ~3 [avatar split clause 255,143,261,219]
581. $false <- (~2, 6, 15) [unit resulting resolution 160,140,220,105] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
585. 2 | ~6 | ~15 [avatar contradiction clause 581]
586. $false [avatar sat refutation 136,141,146,151,161,348,574,577,578,585] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:213,thDist:-213}
% SZS output end Proof for subset_1__t34_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% Introduced skolems: 13
% 
% >>> Saturation
% Initial clauses: 44
% Generated clauses: 403
% Active clauses: 116
% Passive clauses: 209
% Final active clauses: 112
% Final passive clauses: 80
% Main loop iterations started: 129
% 
% >>> Simplifying Inferences
% Duplicate literals: 35
% Trivial inequalities: 1
% Fw subsumption resolutions: 20
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Equational tautologies: 2
% Forward subsumptions: 103
% 
% >>> Generating Inferences
% Binary resolution: 180
% Unit resulting resolution: 4
% Forward superposition: 27
% Backward superposition: 45
% Self superposition: 1
% Equality factoring: 6
% Equality resolution: 6
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 33
% SAT solver unit clauses: 12
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 5245
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% forward demodulation: 0.001 s ( own 0.0000 s ) 
% resolution: 0.001 s
% unit resulting resolution: 0.003 s
% other: 0.005 s ( own 0.0000 s ) 

