i: [1,46,0,3,1,1,0,0] t20_arytm_1
a: [3,69,0,3,1,46]
s: 69
i: [1,45,0,3,1,1,0,0] t20_arytm_1
a: [3,74,0,3,1,45]
s: 74
i: [1,44,0,3,1,1,0,0] t20_arytm_1
a: [3,79,0,3,1,44]
s: 79
i: [1,43,0,3,1,1,0,0] t20_arytm_1
a: [3,84,0,3,1,43]
s: 84
i: [1,47,0,11,1,1,0,0] t20_arytm_1
a: [3,64,0,11,1,47]
s: 64
i: [1,48,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,89,0,3,1,48]
s: 89
i: [1,49,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,94,0,3,1,49]
s: 94
i: [1,55,0,4,1,0,0,1] existence_m1_subset_1
s: 55
i: [1,56,0,11,3,0,0,1] dt_k1_arytm_1
s: 56
i: [1,57,0,11,3,0,0,1] dt_k7_arytm_2
s: 57
i: [1,59,0,12,4,0,0,1] connectedness_r1_arytm_2
s: 59
d: [2,98,1,9,3,1,40,59,84]
s: 98
i: [1,58,0,13,3,0,0,1] commutativity_k7_arytm_2
s: 58
d: [2,104,2,6,2,1,40,98,84]
d: [2,108,2,3,1,1,18,104]
a: [3,112,2,3,1,108]
s: 112
i: [1,50,0,16,4,0,0,255] d2_arytm_1
s: 50
i: [1,60,0,16,4,0,0,1] t19_arytm_2
s: 60
d: [2,106,2,6,2,2,40,98,74]
a: [3,121,2,3,1,106]
s: 121
d: [2,163,0,11,3,0,21,60,57]
s: 163
d: [2,107,2,8,2,1,40,98,55]
a: [3,126,2,4,1,107]
s: 126
i: [1,51,0,18,4,0,0,255] d2_arytm_1
s: 51
i: [1,54,0,18,6,0,0,1] t3_arytm_1
s: 54
d: [2,148,1,24,4,1,44,64,50]
d: [2,149,1,21,3,2,21,148,84]
a: [3,153,1,5,1,149]
s: 153
d: [2,182,2,0,0,3,52,74,79,153,57]
e: 182
d: [2,183,2,6,2,1,40,153,57]
d: [2,187,2,3,1,2,21,183,79]
d: [2,188,2,0,0,3,21,187,74]
e: 188
r: 153
i: [1,52,0,23,5,0,0,1] t13_arytm_1
s: 52
a: [3,152,1,5,1,149]
s: 152
a: [3,156,1,5,1,149]
s: 156
r: 156
a: [3,157,1,5,1,149]
s: 157
d: [2,212,2,10,2,2,40,152,98]
a: [3,219,2,5,1,212]
s: 219
r: 157
r: 219
s: 156
a: [3,218,2,5,1,212]
s: 218
d: [2,99,1,9,3,1,40,59,79]
s: 99
d: [2,100,1,9,3,1,40,59,74]
s: 100
d: [2,282,2,6,2,1,40,99,79]
d: [2,285,2,3,1,1,18,282]
a: [3,298,2,3,1,285]
s: 298
d: [2,168,1,15,5,1,40,54,69]
d: [2,175,1,12,4,2,21,168,79]
d: [2,176,1,9,3,3,21,175,84]
s: 176
d: [2,283,2,6,2,2,40,99,74]
a: [3,303,2,3,1,283]
s: 303
d: [2,132,1,17,3,1,44,64,58]
d: [2,138,1,14,2,2,21,132,79]
d: [2,139,1,11,1,3,21,138,74]
a: [3,143,1,11,1,139]
s: 143
a: [3,161,1,11,1,149]
s: 161
d: [2,370,2,0,0,5,52,84,79,74,69,161,52]
e: 370
d: [2,371,2,23,5,1,44,161,52]
d: [2,372,2,12,4,1,19,371]
d: [2,374,2,9,3,2,21,372,69]
d: [2,375,2,6,2,3,21,374,74]
d: [2,376,2,3,1,4,21,375,79]
d: [2,377,2,0,0,5,21,376,84]
e: 377
r: 156
r: 161
r: 218
d: [2,323,2,6,2,1,40,100,74]
d: [2,325,2,3,1,1,18,323]
a: [3,339,2,3,1,325]
s: 339
s: 157
d: [2,284,2,8,2,1,40,99,55]
a: [3,312,2,4,1,284]
s: 312
a: [3,160,1,11,1,149]
s: 160
r: 143
r: 64
d: [2,393,0,11,1,2,23,64,160]
a: [3,404,0,11,1,393]
s: 404
d: [2,324,2,8,2,1,40,100,55]
a: [3,344,2,4,1,324]
s: 344
d: [2,392,1,11,1,2,23,143,160]
a: [3,399,1,11,1,392]
s: 399
s: 219
d: [2,422,1,17,3,1,44,404,58]
d: [2,426,1,14,2,2,21,422,74]
d: [2,427,1,11,1,3,21,426,79]
a: [3,431,1,11,1,427]
s: 431
d: [2,101,1,11,3,0,40,59,55]
s: 101
d: [2,222,2,11,3,1,44,156,58]
d: [2,226,2,8,2,2,21,222,74]
d: [2,227,2,5,1,3,21,226,79]
a: [3,230,2,5,1,227]
s: 230
d: [2,165,1,17,5,0,44,163,58]
d: [2,166,1,11,3,0,18,165]
s: 166
d: [2,280,2,10,2,2,40,99,152]
a: [3,289,2,5,1,280]
s: 289
d: [2,164,1,17,5,0,44,163,58]
d: [2,167,1,11,3,0,18,164]
s: 167
d: [2,171,1,25,7,0,40,54,163]
d: [2,173,1,22,6,0,18,171]
d: [2,179,1,17,5,0,21,173,57]
s: 179
d: [2,507,2,0,0,5,52,74,84,79,69,157,179]
e: 507
d: [2,508,3,0,0,5,52,79,84,74,121,230,179]
e: 508
d: [2,510,2,12,4,1,40,179,157]
d: [2,523,2,9,3,2,21,510,69]
d: [2,524,2,6,2,3,21,523,79]
d: [2,525,2,3,1,4,21,524,84]
d: [2,526,2,0,0,5,21,525,74]
e: 526
d: [2,511,3,12,4,1,40,179,230]
d: [2,528,3,9,3,2,21,511,121]
d: [2,529,3,6,2,3,21,528,74]
d: [2,530,3,3,1,4,21,529,84]
d: [2,531,3,0,0,5,21,530,79]
e: 531
f: 526,377,188,149,43,44,45,46,47
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t20_arytm_1
% SZS output start Proof for arytm_1__t20_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (r1_arytm_2(X0,X1) => k7_arytm_2(X2,k1_arytm_1(X1,X0)) = k2_arytm_1(k7_arytm_2(X2,X1),X0))))) [input t20_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (r1_arytm_2(X0,X1) => k7_arytm_2(X2,k1_arytm_1(X1,X0)) = k2_arytm_1(k7_arytm_2(X2,X1),X0))))) [negated conjecture 1]
5. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((~r1_arytm_2(X1,X0) => k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,k1_arytm_1(X1,X0))) & (r1_arytm_2(X1,X0) => k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1))))) [input d2_arytm_1]
13. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2)) [input dt_k7_arytm_2]
17. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (r1_arytm_2(X0,X1) => k7_arytm_2(X2,k1_arytm_1(X1,X0)) = k1_arytm_1(k7_arytm_2(X2,X1),X0))))) [input t13_arytm_1]
18. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (k7_arytm_2(X1,X2) = X0 => r1_arytm_2(X2,X0))))) [input t19_arytm_2]
19. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => ((r1_arytm_2(X1,X2) & r1_arytm_2(X0,X1)) => r1_arytm_2(X0,X2))))) [input t3_arytm_1]
20. ? [X0] : (? [X1] : (? [X2] : ((k7_arytm_2(X2,k1_arytm_1(X1,X0)) != k2_arytm_1(k7_arytm_2(X2,X1),X0) & r1_arytm_2(X0,X1)) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
21. ? [X0] : (? [X1] : (? [X2] : (k7_arytm_2(X2,k1_arytm_1(X1,X0)) != k2_arytm_1(k7_arytm_2(X2,X1),X0) & r1_arytm_2(X0,X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 20]
22. ! [X0] : (! [X1] : (((k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,k1_arytm_1(X1,X0)) | r1_arytm_2(X1,X0)) & (k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 5]
23. ! [X0] : (! [X1] : (! [X2] : ((k7_arytm_2(X2,k1_arytm_1(X1,X0)) = k1_arytm_1(k7_arytm_2(X2,X1),X0) | ~r1_arytm_2(X0,X1)) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 17]
24. ! [X0] : (! [X1] : (! [X2] : (k7_arytm_2(X2,k1_arytm_1(X1,X0)) = k1_arytm_1(k7_arytm_2(X2,X1),X0) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 23]
25. ! [X0] : (! [X1] : (! [X2] : ((r1_arytm_2(X2,X0) | k7_arytm_2(X1,X2) != X0) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 18]
26. ! [X0] : (! [X1] : (! [X2] : (r1_arytm_2(X2,X0) | k7_arytm_2(X1,X2) != X0 | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 25]
27. ! [X0] : (! [X1] : (! [X2] : ((r1_arytm_2(X0,X2) | (~r1_arytm_2(X1,X2) | ~r1_arytm_2(X0,X1))) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 19]
28. ! [X0] : (! [X1] : (! [X2] : (r1_arytm_2(X0,X2) | ~r1_arytm_2(X1,X2) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 27]
31. ! [X0,X1] : (m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 13]
32. ! [X0,X1] : (m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 31]
37. ? [X0] : (? [X1] : (? [X2] : (k7_arytm_2(X2,k1_arytm_1(X1,X0)) != k2_arytm_1(k7_arytm_2(X2,X1),X0) & r1_arytm_2(X0,X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (? [X2] : (k7_arytm_2(X2,k1_arytm_1(X1,sK0)) != k2_arytm_1(k7_arytm_2(X2,X1),sK0) & r1_arytm_2(sK0,X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
38. ? [X1] : (? [X2] : (k7_arytm_2(X2,k1_arytm_1(X1,sK0)) != k2_arytm_1(k7_arytm_2(X2,X1),sK0) & r1_arytm_2(sK0,X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) => (? [X2] : (k7_arytm_2(X2,k1_arytm_1(sK1,sK0)) != k2_arytm_1(k7_arytm_2(X2,sK1),sK0) & r1_arytm_2(sK0,sK1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
39. ? [X2] : (k7_arytm_2(X2,k1_arytm_1(sK1,sK0)) != k2_arytm_1(k7_arytm_2(X2,sK1),sK0) & r1_arytm_2(sK0,sK1) & m1_subset_1(X2,k2_arytm_2)) => (k7_arytm_2(sK2,k1_arytm_1(sK1,sK0)) != k2_arytm_1(k7_arytm_2(sK2,sK1),sK0) & r1_arytm_2(sK0,sK1) & m1_subset_1(sK2,k2_arytm_2)) [choice axiom]
40. ((k7_arytm_2(sK2,k1_arytm_1(sK1,sK0)) != k2_arytm_1(k7_arytm_2(sK2,sK1),sK0) & r1_arytm_2(sK0,sK1) & m1_subset_1(sK2,k2_arytm_2)) & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 21,39,38,37]
43. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
44. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
45. m1_subset_1(sK2,k2_arytm_2) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
46. r1_arytm_2(sK0,sK1) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
47. k7_arytm_2(sK2,k1_arytm_1(sK1,sK0)) != k2_arytm_1(k7_arytm_2(sK2,sK1),sK0) [cnf transformation 40] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
50. k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1) | ~r1_arytm_2(X1,X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 22] {a:0,w:16,wCS:6400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
52. k7_arytm_2(X2,k1_arytm_1(X1,X0)) = k1_arytm_1(k7_arytm_2(X2,X1),X0) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 24] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. r1_arytm_2(X2,X0) | k7_arytm_2(X1,X2) != X0 | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 26] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
54. ~r1_arytm_2(X1,X2) | r1_arytm_2(X0,X2) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 28] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 32] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. r1_arytm_2(X2,k7_arytm_2(X1,X2)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k7_arytm_2(X1,X2),k2_arytm_2) [equality resolution 53] {a:0,w:16,wCS:6400,sine:1,thAx:0,allAx:1,thDist:-1}
62. 1 <=> k7_arytm_2(sK2,k1_arytm_1(sK1,sK0)) = k2_arytm_1(k7_arytm_2(sK2,sK1),sK0) [avatar definition]
64. k7_arytm_2(sK2,k1_arytm_1(sK1,sK0)) != k2_arytm_1(k7_arytm_2(sK2,sK1),sK0) <- (~1) [avatar component clause 62] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
65. ~1 [avatar split clause 47,62]
67. 2 <=> r1_arytm_2(sK0,sK1) [avatar definition]
69. r1_arytm_2(sK0,sK1) <- (2) [avatar component clause 67] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
70. 2 [avatar split clause 46,67]
72. 3 <=> m1_subset_1(sK2,k2_arytm_2) [avatar definition]
74. m1_subset_1(sK2,k2_arytm_2) <- (3) [avatar component clause 72] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
75. 3 [avatar split clause 45,72]
77. 4 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
79. m1_subset_1(sK1,k2_arytm_2) <- (4) [avatar component clause 77] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
80. 4 [avatar split clause 44,77]
82. 5 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
84. m1_subset_1(sK0,k2_arytm_2) <- (5) [avatar component clause 82] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
85. 5 [avatar split clause 43,82]
148. k7_arytm_2(sK2,k1_arytm_1(sK1,sK0)) != k1_arytm_1(k7_arytm_2(sK2,sK1),sK0) | ~r1_arytm_2(sK0,k7_arytm_2(sK2,sK1)) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(k7_arytm_2(sK2,sK1),k2_arytm_2) <- (~1) [superposition 64,50] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
149. k7_arytm_2(sK2,k1_arytm_1(sK1,sK0)) != k1_arytm_1(k7_arytm_2(sK2,sK1),sK0) | ~r1_arytm_2(sK0,k7_arytm_2(sK2,sK1)) | ~m1_subset_1(k7_arytm_2(sK2,sK1),k2_arytm_2) <- (~1, 5) [subsumption resolution 148,84] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
151. 14 <=> m1_subset_1(k7_arytm_2(sK2,sK1),k2_arytm_2) [avatar definition]
153. ~m1_subset_1(k7_arytm_2(sK2,sK1),k2_arytm_2) <- (~14) [avatar component clause 151] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
155. 15 <=> r1_arytm_2(sK0,k7_arytm_2(sK2,sK1)) [avatar definition]
157. ~r1_arytm_2(sK0,k7_arytm_2(sK2,sK1)) <- (~15) [avatar component clause 155] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
159. 16 <=> k7_arytm_2(sK2,k1_arytm_1(sK1,sK0)) = k1_arytm_1(k7_arytm_2(sK2,sK1),sK0) [avatar definition]
161. k7_arytm_2(sK2,k1_arytm_1(sK1,sK0)) != k1_arytm_1(k7_arytm_2(sK2,sK1),sK0) <- (~16) [avatar component clause 159] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
162. ~14 | ~15 | ~16 | 1 | ~5 [avatar split clause 149,82,62,159,155,151]
163. r1_arytm_2(X2,k7_arytm_2(X1,X2)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) [subsumption resolution 60,57] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
171. r1_arytm_2(X3,k7_arytm_2(X4,X5)) | ~r1_arytm_2(X3,X5) | ~m1_subset_1(k7_arytm_2(X4,X5),k2_arytm_2) | ~m1_subset_1(X5,k2_arytm_2) | ~m1_subset_1(X3,k2_arytm_2) | ~m1_subset_1(X5,k2_arytm_2) | ~m1_subset_1(X4,k2_arytm_2) [resolution 54,163] {a:1,w:25,wCS:10000,sine:1,thAx:0,allAx:2,thDist:-2}
173. r1_arytm_2(X3,k7_arytm_2(X4,X5)) | ~r1_arytm_2(X3,X5) | ~m1_subset_1(k7_arytm_2(X4,X5),k2_arytm_2) | ~m1_subset_1(X5,k2_arytm_2) | ~m1_subset_1(X3,k2_arytm_2) | ~m1_subset_1(X4,k2_arytm_2) [duplicate literal removal 171] {a:1,w:22,wCS:8800,sine:1,thAx:0,allAx:2,thDist:-2}
179. r1_arytm_2(X3,k7_arytm_2(X4,X5)) | ~r1_arytm_2(X3,X5) | ~m1_subset_1(X5,k2_arytm_2) | ~m1_subset_1(X3,k2_arytm_2) | ~m1_subset_1(X4,k2_arytm_2) [subsumption resolution 173,57] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
183. ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK2,k2_arytm_2) <- (~14) [resolution 153,57] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
187. ~m1_subset_1(sK2,k2_arytm_2) <- (4, ~14) [subsumption resolution 183,79] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
188. $false <- (3, 4, ~14) [subsumption resolution 187,74] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
189. ~3 | ~4 | 14 [avatar contradiction clause 188]
371. k1_arytm_1(k7_arytm_2(sK2,sK1),sK0) != k1_arytm_1(k7_arytm_2(sK2,sK1),sK0) | ~r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~16) [superposition 161,52] {a:2,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
372. ~r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~16) [trivial inequality removal 371] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
374. ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (2, ~16) [subsumption resolution 372,69] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
375. ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (2, 3, ~16) [subsumption resolution 374,74] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
376. ~m1_subset_1(sK0,k2_arytm_2) <- (2, 3, 4, ~16) [subsumption resolution 375,79] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
377. $false <- (2, 3, 4, 5, ~16) [subsumption resolution 376,84] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
378. ~2 | ~3 | ~4 | ~5 | 16 [avatar contradiction clause 377]
510. ~r1_arytm_2(sK0,sK1) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK2,k2_arytm_2) <- (~15) [resolution 179,157] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
523. ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK2,k2_arytm_2) <- (2, ~15) [subsumption resolution 510,69] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
524. ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK2,k2_arytm_2) <- (2, 4, ~15) [subsumption resolution 523,79] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
525. ~m1_subset_1(sK2,k2_arytm_2) <- (2, 4, 5, ~15) [subsumption resolution 524,84] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
526. $false <- (2, 3, 4, 5, ~15) [subsumption resolution 525,74] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
527. ~2 | ~3 | ~4 | ~5 | 15 [avatar contradiction clause 526]
535. $false [avatar sat refutation 65,70,75,80,85,162,189,378,527] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:654,thDist:-654}
% SZS output end Proof for arytm_1__t20_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 302
% Active clauses: 49
% Passive clauses: 89
% Final active clauses: 40
% Final passive clauses: 31
% Main loop iterations started: 51
% 
% >>> Simplifying Inferences
% Duplicate literals: 34
% Trivial inequalities: 3
% Fw subsumption resolutions: 104
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 53
% Unit resulting resolution: 4
% Factoring: 1
% Forward superposition: 39
% Backward superposition: 11
% 
% >>> AVATAR
% Split clauses: 16
% Split components: 35
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 61
% SAT solver unit clauses: 8
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

