i: [1,222,0,2,1,1,0,0] t47_functor0
a: [3,343,0,2,1,222]
s: 343
i: [1,221,0,2,1,1,0,0] t47_functor0
a: [3,348,0,2,1,221]
s: 348
i: [1,220,0,2,1,1,0,0] t47_functor0
a: [3,353,0,2,1,220]
s: 353
i: [1,219,0,2,1,1,0,0] t47_functor0
a: [3,358,0,2,1,219]
s: 358
i: [1,218,0,2,1,1,0,0] t47_functor0
a: [3,363,0,2,1,218]
s: 363
i: [1,217,0,2,1,1,0,0] t47_functor0
a: [3,368,0,2,1,217]
s: 368
i: [1,216,0,2,1,1,0,0] t47_functor0
a: [3,373,0,2,1,216]
s: 373
i: [1,226,0,4,1,1,0,0] t47_functor0
a: [3,323,0,4,1,226]
s: 323
i: [1,225,0,4,1,1,0,0] t47_functor0
a: [3,328,0,4,1,225]
s: 328
i: [1,224,0,4,1,1,0,0] t47_functor0
a: [3,333,0,4,1,224]
s: 333
i: [1,223,0,4,1,1,0,0] t47_functor0
a: [3,338,0,4,1,223]
s: 338
i: [1,227,0,2,1,0,0,255] spc1_boole
a: [3,378,0,2,1,227]
s: 378
i: [1,228,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,383,0,2,1,228]
s: 383
i: [1,287,0,2,1,0,0,2] existence_l1_struct_0
a: [3,393,0,2,1,287]
s: 393
i: [1,288,0,2,1,0,0,3] existence_l1_altcat_1
a: [3,398,0,2,1,288]
s: 398
i: [1,289,0,2,1,0,0,1] existence_l2_altcat_1
a: [3,403,0,2,1,289]
s: 403
i: [1,291,0,2,1,0,0,255] rc7_funct_1
a: [3,408,0,2,1,291]
s: 408
i: [1,290,0,2,1,0,0,255] rc7_funct_1
a: [3,413,0,2,1,290]
s: 413
i: [1,293,0,2,1,0,0,255] rc1_relat_1
a: [3,418,0,2,1,293]
s: 418
i: [1,292,0,2,1,0,0,255] rc1_relat_1
a: [3,423,0,2,1,292]
s: 423
i: [1,296,0,2,1,0,0,255] rc4_funct_1
a: [3,428,0,2,1,296]
s: 428
i: [1,295,0,2,1,0,0,255] rc4_funct_1
a: [3,433,0,2,1,295]
s: 433
i: [1,294,0,2,1,0,0,255] rc4_funct_1
a: [3,438,0,2,1,294]
s: 438
i: [1,298,0,2,1,0,0,255] rc2_altcat_2
a: [3,443,0,2,1,298]
s: 443
i: [1,297,0,2,1,0,0,255] rc2_altcat_2
a: [3,448,0,2,1,297]
s: 448
i: [1,300,0,2,1,0,0,255] rc3_altcat_2
a: [3,453,0,2,1,300]
s: 453
i: [1,299,0,2,1,0,0,255] rc3_altcat_2
a: [3,458,0,2,1,299]
s: 458
i: [1,302,0,2,1,0,0,255] rc7_functor0
a: [3,463,0,2,1,302]
s: 463
i: [1,301,0,2,1,0,0,255] rc7_functor0
a: [3,468,0,2,1,301]
s: 468
i: [1,303,0,2,1,0,0,255] rc2_relat_1
a: [3,473,0,2,1,303]
s: 473
i: [1,305,0,2,1,0,0,255] rc3_funct_1
a: [3,478,0,2,1,305]
s: 478
i: [1,304,0,2,1,0,0,255] rc3_funct_1
a: [3,483,0,2,1,304]
s: 483
i: [1,307,0,2,1,0,0,255] rc1_funct_1
a: [3,488,0,2,1,307]
s: 488
i: [1,306,0,2,1,0,0,255] rc1_funct_1
a: [3,493,0,2,1,306]
s: 493
i: [1,310,0,2,1,0,0,255] rc5_funct_1
a: [3,498,0,2,1,310]
s: 498
i: [1,309,0,2,1,0,0,255] rc5_funct_1
a: [3,503,0,2,1,309]
s: 503
i: [1,308,0,2,1,0,0,255] rc5_funct_1
a: [3,508,0,2,1,308]
s: 508
i: [1,312,0,2,1,0,0,255] rc2_funct_1
a: [3,513,0,2,1,312]
s: 513
i: [1,311,0,2,1,0,0,255] rc2_funct_1
a: [3,518,0,2,1,311]
s: 518
i: [1,315,0,2,1,0,0,255] rc1_funcop_1
a: [3,523,0,2,1,315]
s: 523
i: [1,314,0,2,1,0,0,255] rc1_funcop_1
a: [3,528,0,2,1,314]
s: 528
i: [1,313,0,2,1,0,0,255] rc1_funcop_1
a: [3,533,0,2,1,313]
s: 533
i: [1,319,0,2,1,0,0,255] rc2_funcop_1
a: [3,538,0,2,1,319]
s: 538
i: [1,318,0,2,1,0,0,255] rc2_funcop_1
a: [3,543,0,2,1,318]
s: 543
i: [1,317,0,2,1,0,0,255] rc2_funcop_1
a: [3,548,0,2,1,317]
s: 548
i: [1,316,0,2,1,0,0,255] rc2_funcop_1
a: [3,553,0,2,1,316]
s: 553
i: [1,229,0,3,1,0,0,255] d2_xboole_0
a: [3,388,0,3,1,229]
s: 388
r: 383
d: [2,555,0,2,1,2,23,383,388]
a: [3,559,0,2,1,555]
s: 559
i: [1,232,0,4,2,0,0,255] cc7_funct_1
s: 232
i: [1,233,0,4,2,0,0,255] cc1_funct_1
s: 233
i: [1,234,0,4,2,0,0,255] cc1_relat_1
s: 234
d: [2,561,1,2,1,1,40,232,559]
a: [3,565,1,2,1,561]
s: 565
i: [1,251,0,4,2,0,0,3] dt_l1_altcat_1
s: 251
d: [2,567,1,2,1,1,40,233,559]
a: [3,571,1,2,1,567]
s: 571
i: [1,252,0,4,2,0,0,3] dt_l2_altcat_1
s: 252
i: [1,278,0,4,1,0,0,255] existence_m1_subset_1
s: 278
d: [2,591,1,2,1,1,40,252,358]
a: [3,599,1,2,1,591]
s: 599
i: [1,235,0,5,2,0,0,255] t6_boole
s: 235
d: [2,592,1,2,1,1,40,252,343]
a: [3,604,1,2,1,592]
s: 604
i: [1,277,0,5,2,0,0,255] t7_boole
s: 277
i: [1,236,0,6,3,0,0,255] cc7_struct_0
s: 236
d: [2,621,2,2,1,1,40,599,251]
a: [3,625,2,2,1,621]
s: 625
i: [1,237,0,6,3,0,0,255] cc5_struct_0
s: 237
d: [2,631,2,2,1,1,40,604,251]
a: [3,635,2,2,1,631]
s: 635
i: [1,239,0,6,3,0,0,255] cc17_struct_0
s: 239
i: [1,240,0,6,3,0,0,255] cc2_struct_0
s: 240
d: [2,573,1,2,1,1,40,234,559]
a: [3,577,1,2,1,573]
s: 577
i: [1,241,0,6,3,0,0,255] cc6_struct_0
s: 241
d: [2,579,1,2,1,1,40,251,398]
a: [3,584,1,2,1,579]
s: 584
i: [1,242,0,6,3,0,0,255] cc1_struct_0
s: 242
i: [1,245,0,6,3,0,0,255] cc4_struct_0
s: 245
d: [2,580,1,2,1,1,40,251,448]
a: [3,589,1,2,1,580]
s: 589
i: [1,253,0,6,3,0,0,255] cc4_altcat_1
s: 253
d: [2,593,1,2,1,1,40,252,403]
a: [3,609,1,2,1,593]
s: 609
i: [1,274,0,6,2,0,0,255] t1_subset
s: 274
i: [1,279,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 279
d: [2,594,1,2,1,1,40,252,458]
a: [3,614,1,2,1,594]
s: 614
i: [1,230,0,7,3,0,0,255] cc8_funct_1
s: 230
d: [2,595,1,2,1,1,40,252,468]
a: [3,619,1,2,1,595]
s: 619
i: [1,231,0,7,3,0,0,255] cc8_funct_1
s: 231
i: [1,243,0,7,3,0,0,255] cc8_struct_0
s: 243
d: [2,637,1,4,2,1,40,237,393]
a: [3,642,1,2,1,637]
s: 642
i: [1,246,0,7,3,0,0,255] cc9_struct_0
s: 246
d: [2,658,2,4,2,1,40,584,239]
a: [3,663,2,2,1,658]
s: 663
i: [1,247,0,7,3,0,0,255] cc11_struct_0
s: 247
i: [1,248,0,7,3,0,0,255] cc11_struct_0
s: 248
d: [2,676,2,2,1,1,40,609,251]
a: [3,680,2,2,1,676]
s: 680
i: [1,254,0,7,3,0,0,255] cc3_altcat_1
s: 254
d: [2,682,2,2,1,1,40,614,251]
a: [3,686,2,2,1,682]
s: 686
i: [1,276,0,7,3,0,0,255] t8_boole
s: 276
i: [1,264,0,8,4,0,0,255] cc4_funct_1
s: 264
d: [2,689,2,2,1,1,40,619,251]
a: [3,693,2,2,1,689]
s: 693
d: [2,718,0,6,3,0,21,264,233]
d: [2,719,0,4,2,0,21,718,234]
s: 719
d: [2,722,1,2,1,1,40,719,559]
a: [3,726,1,2,1,722]
s: 726
i: [1,273,0,8,4,0,0,255] cc3_funcop_1
s: 273
d: [2,729,0,6,3,0,21,273,233]
d: [2,730,0,4,2,0,21,729,234]
s: 730
d: [2,731,1,2,1,1,40,730,559]
a: [3,735,1,2,1,731]
s: 735
i: [1,275,0,8,3,0,0,255] t2_subset
s: 275
d: [2,700,3,4,2,1,40,680,239]
a: [3,705,3,2,1,700]
s: 705
i: [1,249,0,9,4,0,0,255] cc10_struct_0
s: 249
i: [1,282,0,10,3,0,0,2] existence_l1_functor0
s: 282
d: [2,688,1,5,2,0,40,230,278]
s: 688
i: [1,284,0,10,3,0,0,3] existence_l2_functor0
s: 284
d: [2,695,1,5,2,0,40,231,278]
s: 695
i: [1,285,0,10,3,0,0,2] rc5_functor0
s: 285
i: [1,286,0,10,3,0,0,2] rc5_functor0
s: 286
d: [2,738,1,6,2,0,40,275,278]
s: 738
i: [1,283,0,12,4,0,0,2] dt_l2_functor0
s: 283
d: [2,744,2,6,2,0,40,738,279]
s: 744
i: [1,250,0,20,6,0,0,1] t45_functor0
s: 250
i: [1,281,0,20,8,0,0,1] existence_m2_functor0
s: 281
d: [2,627,1,7,3,0,44,235,235]
s: 627
i: [1,280,0,22,9,0,0,3] dt_m2_functor0
s: 280
d: [2,754,1,18,8,1,40,280,338]
d: [2,757,1,16,7,2,21,754,343]
d: [2,758,1,14,6,3,21,757,348]
d: [2,759,1,12,5,4,21,758,353]
d: [2,760,1,10,4,5,21,759,358]
d: [2,761,1,8,3,6,21,760,363]
d: [2,762,1,6,2,7,21,761,368]
d: [2,763,1,4,1,8,21,762,373]
a: [3,767,1,4,1,763]
s: 767
i: [1,257,0,26,10,0,0,255] d27_functor0
s: 257
i: [1,258,0,26,10,0,0,255] d27_functor0
s: 258
d: [2,777,1,22,9,1,40,257,333]
d: [2,778,1,18,8,2,21,777,338]
d: [2,779,1,16,7,3,21,778,343]
d: [2,780,1,14,6,4,21,779,348]
d: [2,781,1,12,5,5,21,780,353]
d: [2,782,1,10,4,6,21,781,358]
d: [2,783,1,8,3,7,21,782,363]
d: [2,784,1,6,2,8,21,783,368]
d: [2,785,1,4,1,9,21,784,373]
a: [3,789,1,4,1,785]
s: 789
d: [2,769,2,8,3,1,40,767,283]
d: [2,770,2,6,2,2,21,769,604]
d: [2,771,2,4,1,3,21,770,599]
d: [2,805,2,16,5,1,40,789,250]
d: [2,806,2,12,4,2,21,805,328]
d: [2,807,2,8,3,3,21,806,323]
d: [2,808,2,6,2,4,21,807,635]
d: [2,809,2,4,1,5,21,808,625]
f: 809,785,771,763,631,621,592,591,216,217,218,219,220,221,222,223,224,225,226
% Refutation found. Thanks to Tanya!
% SZS status Theorem for functor0__t47_functor0
% SZS output start Proof for functor0__t47_functor0
1. ! [X0] : ((l2_altcat_1(X0) & v12_altcat_1(X0) & v2_altcat_1(X0) & ~v2_struct_0(X0)) => ! [X1] : ((l2_altcat_1(X1) & v12_altcat_1(X1) & ~v2_struct_0(X1)) => ! [X2] : (m2_functor0(X2,X0,X1) => ((v5_functor0(X2,X0,X1) & v16_functor0(X2,X0,X1)) => v7_functor0(X2,X0,X1))))) [input t47_functor0]
2. ~! [X0] : ((l2_altcat_1(X0) & v12_altcat_1(X0) & v2_altcat_1(X0) & ~v2_struct_0(X0)) => ! [X1] : ((l2_altcat_1(X1) & v12_altcat_1(X1) & ~v2_struct_0(X1)) => ! [X2] : (m2_functor0(X2,X0,X1) => ((v5_functor0(X2,X0,X1) & v16_functor0(X2,X0,X1)) => v7_functor0(X2,X0,X1))))) [negated conjecture 1]
33. ! [X0] : ((l2_altcat_1(X0) & v12_altcat_1(X0) & v2_altcat_1(X0) & ~v2_struct_0(X0)) => ! [X1] : ((l2_altcat_1(X1) & v12_altcat_1(X1) & ~v2_struct_0(X1)) => ! [X2] : (m2_functor0(X2,X0,X1) => (v16_functor0(X2,X0,X1) <=> (v14_functor0(X2,X0,X1) & v11_functor0(X2,X0,X1)))))) [input d27_functor0]
36. ! [X0] : (l1_altcat_1(X0) => l1_struct_0(X0)) [input dt_l1_altcat_1]
39. ! [X0] : (l2_altcat_1(X0) => l1_altcat_1(X0)) [input dt_l2_altcat_1]
40. ! [X0,X1] : ((l1_altcat_1(X1) & l1_altcat_1(X0)) => ! [X2] : (l2_functor0(X2,X0,X1) => l1_functor0(X2,X0,X1))) [input dt_l2_functor0]
42. ! [X0,X1] : ((l2_altcat_1(X1) & v12_altcat_1(X1) & ~v2_struct_0(X1) & l2_altcat_1(X0) & v12_altcat_1(X0) & v2_altcat_1(X0) & ~v2_struct_0(X0)) => ! [X2] : (m2_functor0(X2,X0,X1) => l2_functor0(X2,X0,X1))) [input dt_m2_functor0]
68. ! [X0] : (l1_struct_0(X0) => ! [X1] : (l1_struct_0(X1) => ! [X2] : (l1_functor0(X2,X0,X1) => ((v5_functor0(X2,X0,X1) & v11_functor0(X2,X0,X1)) => v7_functor0(X2,X0,X1))))) [input t45_functor0]
93. ? [X0] : (? [X1] : (? [X2] : ((~v7_functor0(X2,X0,X1) & (v5_functor0(X2,X0,X1) & v16_functor0(X2,X0,X1))) & m2_functor0(X2,X0,X1)) & (l2_altcat_1(X1) & v12_altcat_1(X1) & ~v2_struct_0(X1))) & (l2_altcat_1(X0) & v12_altcat_1(X0) & v2_altcat_1(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
94. ? [X0] : (? [X1] : (? [X2] : (~v7_functor0(X2,X0,X1) & v5_functor0(X2,X0,X1) & v16_functor0(X2,X0,X1) & m2_functor0(X2,X0,X1)) & l2_altcat_1(X1) & v12_altcat_1(X1) & ~v2_struct_0(X1)) & l2_altcat_1(X0) & v12_altcat_1(X0) & v2_altcat_1(X0) & ~v2_struct_0(X0)) [flattening 93]
122. ! [X0] : (! [X1] : (! [X2] : ((v7_functor0(X2,X0,X1) | (~v5_functor0(X2,X0,X1) | ~v11_functor0(X2,X0,X1))) | ~l1_functor0(X2,X0,X1)) | ~l1_struct_0(X1)) | ~l1_struct_0(X0)) [ennf transformation 68]
123. ! [X0] : (! [X1] : (! [X2] : (v7_functor0(X2,X0,X1) | ~v5_functor0(X2,X0,X1) | ~v11_functor0(X2,X0,X1) | ~l1_functor0(X2,X0,X1)) | ~l1_struct_0(X1)) | ~l1_struct_0(X0)) [flattening 122]
124. ! [X0] : (l1_struct_0(X0) | ~l1_altcat_1(X0)) [ennf transformation 36]
125. ! [X0] : (l1_altcat_1(X0) | ~l2_altcat_1(X0)) [ennf transformation 39]
134. ! [X0] : (! [X1] : (! [X2] : ((v16_functor0(X2,X0,X1) <=> (v14_functor0(X2,X0,X1) & v11_functor0(X2,X0,X1))) | ~m2_functor0(X2,X0,X1)) | (~l2_altcat_1(X1) | ~v12_altcat_1(X1) | v2_struct_0(X1))) | (~l2_altcat_1(X0) | ~v12_altcat_1(X0) | ~v2_altcat_1(X0) | v2_struct_0(X0))) [ennf transformation 33]
135. ! [X0] : (! [X1] : (! [X2] : ((v16_functor0(X2,X0,X1) <=> (v14_functor0(X2,X0,X1) & v11_functor0(X2,X0,X1))) | ~m2_functor0(X2,X0,X1)) | ~l2_altcat_1(X1) | ~v12_altcat_1(X1) | v2_struct_0(X1)) | ~l2_altcat_1(X0) | ~v12_altcat_1(X0) | ~v2_altcat_1(X0) | v2_struct_0(X0)) [flattening 134]
156. ! [X0,X1] : (! [X2] : (l2_functor0(X2,X0,X1) | ~m2_functor0(X2,X0,X1)) | (~l2_altcat_1(X1) | ~v12_altcat_1(X1) | v2_struct_0(X1) | ~l2_altcat_1(X0) | ~v12_altcat_1(X0) | ~v2_altcat_1(X0) | v2_struct_0(X0))) [ennf transformation 42]
157. ! [X0,X1] : (! [X2] : (l2_functor0(X2,X0,X1) | ~m2_functor0(X2,X0,X1)) | ~l2_altcat_1(X1) | ~v12_altcat_1(X1) | v2_struct_0(X1) | ~l2_altcat_1(X0) | ~v12_altcat_1(X0) | ~v2_altcat_1(X0) | v2_struct_0(X0)) [flattening 156]
162. ! [X0,X1] : (! [X2] : (l1_functor0(X2,X0,X1) | ~l2_functor0(X2,X0,X1)) | (~l1_altcat_1(X1) | ~l1_altcat_1(X0))) [ennf transformation 40]
163. ! [X0,X1] : (! [X2] : (l1_functor0(X2,X0,X1) | ~l2_functor0(X2,X0,X1)) | ~l1_altcat_1(X1) | ~l1_altcat_1(X0)) [flattening 162]
168. ? [X0] : (? [X1] : (? [X2] : (~v7_functor0(X2,X0,X1) & v5_functor0(X2,X0,X1) & v16_functor0(X2,X0,X1) & m2_functor0(X2,X0,X1)) & l2_altcat_1(X1) & v12_altcat_1(X1) & ~v2_struct_0(X1)) & l2_altcat_1(X0) & v12_altcat_1(X0) & v2_altcat_1(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (~v7_functor0(X2,sK0,X1) & v5_functor0(X2,sK0,X1) & v16_functor0(X2,sK0,X1) & m2_functor0(X2,sK0,X1)) & l2_altcat_1(X1) & v12_altcat_1(X1) & ~v2_struct_0(X1)) & l2_altcat_1(sK0) & v12_altcat_1(sK0) & v2_altcat_1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
169. ? [X1] : (? [X2] : (~v7_functor0(X2,sK0,X1) & v5_functor0(X2,sK0,X1) & v16_functor0(X2,sK0,X1) & m2_functor0(X2,sK0,X1)) & l2_altcat_1(X1) & v12_altcat_1(X1) & ~v2_struct_0(X1)) => (? [X2] : (~v7_functor0(X2,sK0,sK1) & v5_functor0(X2,sK0,sK1) & v16_functor0(X2,sK0,sK1) & m2_functor0(X2,sK0,sK1)) & l2_altcat_1(sK1) & v12_altcat_1(sK1) & ~v2_struct_0(sK1)) [choice axiom]
170. ? [X2] : (~v7_functor0(X2,sK0,sK1) & v5_functor0(X2,sK0,sK1) & v16_functor0(X2,sK0,sK1) & m2_functor0(X2,sK0,sK1)) => (~v7_functor0(sK2,sK0,sK1) & v5_functor0(sK2,sK0,sK1) & v16_functor0(sK2,sK0,sK1) & m2_functor0(sK2,sK0,sK1)) [choice axiom]
171. ((~v7_functor0(sK2,sK0,sK1) & v5_functor0(sK2,sK0,sK1) & v16_functor0(sK2,sK0,sK1) & m2_functor0(sK2,sK0,sK1)) & l2_altcat_1(sK1) & v12_altcat_1(sK1) & ~v2_struct_0(sK1)) & l2_altcat_1(sK0) & v12_altcat_1(sK0) & v2_altcat_1(sK0) & ~v2_struct_0(sK0) [skolemisation 94,170,169,168]
172. ! [X0] : (! [X1] : (! [X2] : (((v16_functor0(X2,X0,X1) | (~v14_functor0(X2,X0,X1) | ~v11_functor0(X2,X0,X1))) & ((v14_functor0(X2,X0,X1) & v11_functor0(X2,X0,X1)) | ~v16_functor0(X2,X0,X1))) | ~m2_functor0(X2,X0,X1)) | ~l2_altcat_1(X1) | ~v12_altcat_1(X1) | v2_struct_0(X1)) | ~l2_altcat_1(X0) | ~v12_altcat_1(X0) | ~v2_altcat_1(X0) | v2_struct_0(X0)) [nnf transformation 135]
173. ! [X0] : (! [X1] : (! [X2] : (((v16_functor0(X2,X0,X1) | ~v14_functor0(X2,X0,X1) | ~v11_functor0(X2,X0,X1)) & ((v14_functor0(X2,X0,X1) & v11_functor0(X2,X0,X1)) | ~v16_functor0(X2,X0,X1))) | ~m2_functor0(X2,X0,X1)) | ~l2_altcat_1(X1) | ~v12_altcat_1(X1) | v2_struct_0(X1)) | ~l2_altcat_1(X0) | ~v12_altcat_1(X0) | ~v2_altcat_1(X0) | v2_struct_0(X0)) [flattening 172]
216. ~v2_struct_0(sK0) [cnf transformation 171] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
217. v2_altcat_1(sK0) [cnf transformation 171] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
218. v12_altcat_1(sK0) [cnf transformation 171] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
219. l2_altcat_1(sK0) [cnf transformation 171] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
220. ~v2_struct_0(sK1) [cnf transformation 171] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
221. v12_altcat_1(sK1) [cnf transformation 171] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
222. l2_altcat_1(sK1) [cnf transformation 171] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
223. m2_functor0(sK2,sK0,sK1) [cnf transformation 171] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
224. v16_functor0(sK2,sK0,sK1) [cnf transformation 171] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
225. v5_functor0(sK2,sK0,sK1) [cnf transformation 171] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
226. ~v7_functor0(sK2,sK0,sK1) [cnf transformation 171] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
250. ~v11_functor0(X2,X0,X1) | ~v5_functor0(X2,X0,X1) | v7_functor0(X2,X0,X1) | ~l1_functor0(X2,X0,X1) | ~l1_struct_0(X1) | ~l1_struct_0(X0) [cnf transformation 123] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
251. ~l1_altcat_1(X0) | l1_struct_0(X0) [cnf transformation 124] {a:0,w:4,wCS:1600,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
252. ~l2_altcat_1(X0) | l1_altcat_1(X0) [cnf transformation 125] {a:0,w:4,wCS:1600,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
257. ~v16_functor0(X2,X0,X1) | v11_functor0(X2,X0,X1) | ~m2_functor0(X2,X0,X1) | ~l2_altcat_1(X1) | ~v12_altcat_1(X1) | v2_struct_0(X1) | ~l2_altcat_1(X0) | ~v12_altcat_1(X0) | ~v2_altcat_1(X0) | v2_struct_0(X0) [cnf transformation 173] {a:0,w:26,wCS:10400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
280. ~m2_functor0(X2,X0,X1) | l2_functor0(X2,X0,X1) | ~l2_altcat_1(X1) | ~v12_altcat_1(X1) | v2_struct_0(X1) | ~l2_altcat_1(X0) | ~v12_altcat_1(X0) | ~v2_altcat_1(X0) | v2_struct_0(X0) [cnf transformation 157] {a:0,w:22,wCS:8800,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
283. ~l2_functor0(X2,X0,X1) | l1_functor0(X2,X0,X1) | ~l1_altcat_1(X1) | ~l1_altcat_1(X0) [cnf transformation 163] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
321. 1 <=> v7_functor0(sK2,sK0,sK1) [avatar definition]
323. ~v7_functor0(sK2,sK0,sK1) <- (~1) [avatar component clause 321] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
324. ~1 [avatar split clause 226,321]
326. 2 <=> v5_functor0(sK2,sK0,sK1) [avatar definition]
328. v5_functor0(sK2,sK0,sK1) <- (2) [avatar component clause 326] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
329. 2 [avatar split clause 225,326]
331. 3 <=> v16_functor0(sK2,sK0,sK1) [avatar definition]
333. v16_functor0(sK2,sK0,sK1) <- (3) [avatar component clause 331] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
334. 3 [avatar split clause 224,331]
336. 4 <=> m2_functor0(sK2,sK0,sK1) [avatar definition]
338. m2_functor0(sK2,sK0,sK1) <- (4) [avatar component clause 336] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
339. 4 [avatar split clause 223,336]
341. 5 <=> l2_altcat_1(sK1) [avatar definition]
343. l2_altcat_1(sK1) <- (5) [avatar component clause 341] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
344. 5 [avatar split clause 222,341]
346. 6 <=> v12_altcat_1(sK1) [avatar definition]
348. v12_altcat_1(sK1) <- (6) [avatar component clause 346] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
349. 6 [avatar split clause 221,346]
351. 7 <=> v2_struct_0(sK1) [avatar definition]
353. ~v2_struct_0(sK1) <- (~7) [avatar component clause 351] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
354. ~7 [avatar split clause 220,351]
356. 8 <=> l2_altcat_1(sK0) [avatar definition]
358. l2_altcat_1(sK0) <- (8) [avatar component clause 356] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
359. 8 [avatar split clause 219,356]
361. 9 <=> v12_altcat_1(sK0) [avatar definition]
363. v12_altcat_1(sK0) <- (9) [avatar component clause 361] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
364. 9 [avatar split clause 218,361]
366. 10 <=> v2_altcat_1(sK0) [avatar definition]
368. v2_altcat_1(sK0) <- (10) [avatar component clause 366] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
369. 10 [avatar split clause 217,366]
371. 11 <=> v2_struct_0(sK0) [avatar definition]
373. ~v2_struct_0(sK0) <- (~11) [avatar component clause 371] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
374. ~11 [avatar split clause 216,371]
591. l1_altcat_1(sK0) <- (8) [resolution 252,358] {a:1,w:2,wCS:200,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
592. l1_altcat_1(sK1) <- (5) [resolution 252,343] {a:1,w:2,wCS:200,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
597. 54 <=> l1_altcat_1(sK0) [avatar definition]
599. l1_altcat_1(sK0) <- (54) [avatar component clause 597] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
600. 54 | ~8 [avatar split clause 591,356,597]
602. 55 <=> l1_altcat_1(sK1) [avatar definition]
604. l1_altcat_1(sK1) <- (55) [avatar component clause 602] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
605. 55 | ~5 [avatar split clause 592,341,602]
621. l1_struct_0(sK0) <- (54) [resolution 599,251] {a:2,w:2,wCS:200,goal:1,sine:3,thAx:0,allAx:6,thDist:-6}
623. 59 <=> l1_struct_0(sK0) [avatar definition]
625. l1_struct_0(sK0) <- (59) [avatar component clause 623] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
626. 59 | ~54 [avatar split clause 621,597,623]
631. l1_struct_0(sK1) <- (55) [resolution 604,251] {a:2,w:2,wCS:200,goal:1,sine:3,thAx:0,allAx:6,thDist:-6}
633. 60 <=> l1_struct_0(sK1) [avatar definition]
635. l1_struct_0(sK1) <- (60) [avatar component clause 633] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
636. 60 | ~55 [avatar split clause 631,602,633]
754. l2_functor0(sK2,sK0,sK1) | ~l2_altcat_1(sK1) | ~v12_altcat_1(sK1) | v2_struct_0(sK1) | ~l2_altcat_1(sK0) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (4) [resolution 280,338] {a:1,w:18,wCS:1800,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
757. l2_functor0(sK2,sK0,sK1) | ~v12_altcat_1(sK1) | v2_struct_0(sK1) | ~l2_altcat_1(sK0) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (4, 5) [subsumption resolution 754,343] {a:1,w:16,wCS:1600,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
758. l2_functor0(sK2,sK0,sK1) | v2_struct_0(sK1) | ~l2_altcat_1(sK0) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (4, 5, 6) [subsumption resolution 757,348] {a:1,w:14,wCS:1400,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
759. l2_functor0(sK2,sK0,sK1) | ~l2_altcat_1(sK0) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (4, 5, 6, ~7) [subsumption resolution 758,353] {a:1,w:12,wCS:1200,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
760. l2_functor0(sK2,sK0,sK1) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (4, 5, 6, ~7, 8) [subsumption resolution 759,358] {a:1,w:10,wCS:1000,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
761. l2_functor0(sK2,sK0,sK1) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (4, 5, 6, ~7, 8, 9) [subsumption resolution 760,363] {a:1,w:8,wCS:800,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
762. l2_functor0(sK2,sK0,sK1) | v2_struct_0(sK0) <- (4, 5, 6, ~7, 8, 9, 10) [subsumption resolution 761,368] {a:1,w:6,wCS:600,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
763. l2_functor0(sK2,sK0,sK1) <- (4, 5, 6, ~7, 8, 9, 10, ~11) [subsumption resolution 762,373] {a:1,w:4,wCS:400,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
765. 75 <=> l2_functor0(sK2,sK0,sK1) [avatar definition]
767. l2_functor0(sK2,sK0,sK1) <- (75) [avatar component clause 765] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
768. 75 | ~4 | ~5 | ~6 | 7 | ~8 | ~9 | ~10 | 11 [avatar split clause 763,371,366,361,356,351,346,341,336,765]
769. l1_functor0(sK2,sK0,sK1) | ~l1_altcat_1(sK1) | ~l1_altcat_1(sK0) <- (75) [resolution 767,283] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
770. l1_functor0(sK2,sK0,sK1) | ~l1_altcat_1(sK0) <- (55, 75) [subsumption resolution 769,604] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
771. l1_functor0(sK2,sK0,sK1) <- (54, 55, 75) [subsumption resolution 770,599] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
773. 76 <=> l1_functor0(sK2,sK0,sK1) [avatar definition]
776. 76 | ~54 | ~55 | ~75 [avatar split clause 771,765,602,597,773]
777. v11_functor0(sK2,sK0,sK1) | ~m2_functor0(sK2,sK0,sK1) | ~l2_altcat_1(sK1) | ~v12_altcat_1(sK1) | v2_struct_0(sK1) | ~l2_altcat_1(sK0) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (3) [resolution 257,333] {a:1,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
778. v11_functor0(sK2,sK0,sK1) | ~l2_altcat_1(sK1) | ~v12_altcat_1(sK1) | v2_struct_0(sK1) | ~l2_altcat_1(sK0) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (3, 4) [subsumption resolution 777,338] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
779. v11_functor0(sK2,sK0,sK1) | ~v12_altcat_1(sK1) | v2_struct_0(sK1) | ~l2_altcat_1(sK0) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (3, 4, 5) [subsumption resolution 778,343] {a:1,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
780. v11_functor0(sK2,sK0,sK1) | v2_struct_0(sK1) | ~l2_altcat_1(sK0) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (3, 4, 5, 6) [subsumption resolution 779,348] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
781. v11_functor0(sK2,sK0,sK1) | ~l2_altcat_1(sK0) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (3, 4, 5, 6, ~7) [subsumption resolution 780,353] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
782. v11_functor0(sK2,sK0,sK1) | ~v12_altcat_1(sK0) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (3, 4, 5, 6, ~7, 8) [subsumption resolution 781,358] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
783. v11_functor0(sK2,sK0,sK1) | ~v2_altcat_1(sK0) | v2_struct_0(sK0) <- (3, 4, 5, 6, ~7, 8, 9) [subsumption resolution 782,363] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
784. v11_functor0(sK2,sK0,sK1) | v2_struct_0(sK0) <- (3, 4, 5, 6, ~7, 8, 9, 10) [subsumption resolution 783,368] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
785. v11_functor0(sK2,sK0,sK1) <- (3, 4, 5, 6, ~7, 8, 9, 10, ~11) [subsumption resolution 784,373] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
787. 77 <=> v11_functor0(sK2,sK0,sK1) [avatar definition]
789. v11_functor0(sK2,sK0,sK1) <- (77) [avatar component clause 787] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
790. 77 | ~3 | ~4 | ~5 | ~6 | 7 | ~8 | ~9 | ~10 | 11 [avatar split clause 785,371,366,361,356,351,346,341,336,331,787]
805. ~v5_functor0(sK2,sK0,sK1) | v7_functor0(sK2,sK0,sK1) | ~l1_functor0(sK2,sK0,sK1) | ~l1_struct_0(sK1) | ~l1_struct_0(sK0) <- (77) [resolution 789,250] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
806. v7_functor0(sK2,sK0,sK1) | ~l1_functor0(sK2,sK0,sK1) | ~l1_struct_0(sK1) | ~l1_struct_0(sK0) <- (2, 77) [subsumption resolution 805,328] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
807. ~l1_functor0(sK2,sK0,sK1) | ~l1_struct_0(sK1) | ~l1_struct_0(sK0) <- (~1, 2, 77) [subsumption resolution 806,323] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
808. ~l1_functor0(sK2,sK0,sK1) | ~l1_struct_0(sK0) <- (~1, 2, 60, 77) [subsumption resolution 807,635] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
809. ~l1_functor0(sK2,sK0,sK1) <- (~1, 2, 59, 60, 77) [subsumption resolution 808,625] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
810. ~76 | 1 | ~2 | ~59 | ~60 | ~77 [avatar split clause 809,787,633,623,326,321,773]
811. $false [avatar sat refutation 324,329,334,339,344,349,354,359,364,369,374,600,605,626,636,768,776,790,810] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:365,thDist:-365}
% SZS output end Proof for functor0__t47_functor0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 70
% >>> Preprocessing
% Introduced skolems: 24
% Pure predicates: 7
% 
% >>> Saturation
% Initial clauses: 104
% Generated clauses: 279
% Active clauses: 110
% Passive clauses: 127
% Final active clauses: 109
% Final passive clauses: 10
% Main loop iterations started: 117
% 
% >>> Simplifying Inferences
% Duplicate literals: 14
% Fw subsumption resolutions: 35
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 21
% Forward subsumptions: 21
% 
% >>> Generating Inferences
% Binary resolution: 52
% Forward superposition: 5
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 78
% SAT solver unit clauses: 47
% SAT solver binary clauses: 17
% 
% Memory used [KB]: 9850
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward demodulation: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

