i: [1,109,0,2,1,1,0,0] t62_ordinal3
a: [3,171,0,2,1,109]
s: 171
i: [1,108,0,2,1,1,0,0] t62_ordinal3
a: [3,176,0,2,1,108]
s: 176
i: [1,110,0,7,1,1,0,0] t62_ordinal3
a: [3,166,0,7,1,110]
s: 166
i: [1,111,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,181,0,2,1,111]
s: 181
i: [1,112,0,2,1,0,0,2] fc1_xboole_0
a: [3,186,0,2,1,112]
s: 186
i: [1,159,0,14,5,0,0,1] d5_ordinal3
a: [3,196,0,2,1,159]
s: 196
i: [1,143,0,7,3,0,0,1] reflexivity_r1_ordinal1
a: [3,203,0,2,1,143]
s: 203
r: 196
d: [2,260,1,0,0,2,40,203,176]
e: 260
d: [2,259,1,0,0,2,40,203,171]
e: 259
d: [2,258,1,0,0,2,52,171,203]
e: 258
r: 203
s: 196
i: [1,149,0,2,1,0,0,3] rc2_xboole_0
a: [3,211,0,2,1,149]
s: 211
i: [1,150,0,2,1,0,0,2] rc4_ordinal1
a: [3,216,0,2,1,150]
s: 216
i: [1,151,0,2,1,0,0,3] rc1_xboole_0
a: [3,221,0,2,1,151]
s: 221
i: [1,152,0,2,1,0,0,1] rc1_ordinal1
a: [3,226,0,2,1,152]
s: 226
i: [1,156,0,2,1,0,0,255] rc2_ordinal1
a: [3,231,0,2,1,156]
s: 231
i: [1,155,0,2,1,0,0,255] rc2_ordinal1
a: [3,236,0,2,1,155]
s: 236
i: [1,154,0,2,1,0,0,255] rc2_ordinal1
a: [3,241,0,2,1,154]
s: 241
i: [1,153,0,2,1,0,0,255] rc2_ordinal1
a: [3,246,0,2,1,153]
s: 246
i: [1,158,0,2,1,0,0,2] rc5_ordinal1
a: [3,251,0,2,1,158]
s: 251
i: [1,157,0,2,1,0,0,2] rc5_ordinal1
a: [3,256,0,2,1,157]
s: 256
i: [1,139,0,3,1,0,0,255] reflexivity_r1_tarski
s: 139
i: [1,113,0,3,1,0,0,255] d2_xboole_0
a: [3,191,0,3,1,113]
s: 191
r: 181
i: [1,115,0,4,2,0,0,2] cc6_ordinal1
s: 115
d: [2,265,1,2,1,1,40,115,216]
a: [3,270,1,2,1,265]
s: 270
i: [1,116,0,4,2,0,0,255] cc9_ordinal1
s: 116
d: [2,266,1,2,1,1,40,115,251]
a: [3,275,1,2,1,266]
s: 275
i: [1,117,0,4,2,0,0,2] cc7_ordinal1
s: 117
i: [1,118,0,4,2,0,0,3] cc3_ordinal1
s: 118
d: [2,301,1,0,0,2,52,196,186,118]
e: 301
d: [2,303,1,2,1,1,40,118,186]
d: [2,305,1,0,0,2,21,303,196]
e: 305
r: 196
a: [3,195,0,2,1,159]
s: 195
i: [1,121,0,4,2,0,0,255] cc1_ordinal1
s: 121
d: [2,314,1,2,1,1,40,121,176]
a: [3,328,1,2,1,314]
s: 328
i: [1,122,0,4,2,0,0,255] cc1_ordinal1
s: 122
i: [1,138,0,4,1,0,0,255] existence_m1_subset_1
s: 138
d: [2,315,1,2,1,1,40,121,171]
a: [3,333,1,2,1,315]
s: 333
i: [1,119,0,5,2,0,0,2] t6_boole
s: 119
d: [2,351,1,2,1,1,40,122,176]
a: [3,365,1,2,1,351]
s: 365
i: [1,135,0,5,2,0,0,255] t7_boole
s: 135
a: [3,206,0,5,2,143]
s: 206
d: [2,352,1,2,1,1,40,122,171]
a: [3,370,1,2,1,352]
s: 370
i: [1,129,0,6,3,0,0,255] cc2_ordinal1
s: 129
d: [2,277,1,2,1,1,40,116,221]
a: [3,282,1,2,1,277]
s: 282
i: [1,130,0,6,2,0,0,255] t1_subset
s: 130
i: [1,140,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 140
d: [2,278,1,2,1,1,40,116,186]
a: [3,287,1,2,1,278]
s: 287
i: [1,120,0,7,2,0,0,1] t27_ordinal2
s: 120
d: [2,289,1,2,1,1,40,117,221]
a: [3,294,1,2,1,289]
s: 294
i: [1,127,0,7,3,0,0,255] cc5_ordinal1
s: 127
i: [1,128,0,7,3,0,0,255] cc2_ordinal2
s: 128
d: [2,290,1,2,1,1,40,117,186]
a: [3,299,1,2,1,290]
s: 299
i: [1,132,0,7,2,0,0,255] t3_subset
s: 132
d: [2,302,1,2,1,1,40,118,221]
a: [3,311,1,2,1,302]
s: 311
i: [1,133,0,7,2,0,0,255] t3_subset
s: 133
i: [1,134,0,7,3,0,0,2] t8_boole
s: 134
d: [2,313,1,2,1,1,40,121,195]
a: [3,323,1,2,1,313]
s: 323
i: [1,114,0,8,3,0,0,255] cc10_ordinal1
s: 114
d: [2,317,1,2,1,1,40,121,226]
a: [3,343,1,2,1,317]
s: 343
i: [1,131,0,8,3,0,0,255] t2_subset
s: 131
i: [1,141,0,8,3,0,0,1] fc5_ordinal2
s: 141
d: [2,350,1,2,1,1,40,122,195]
a: [3,360,1,2,1,350]
s: 360
i: [1,142,0,8,3,0,0,1] fc5_ordinal2
s: 142
d: [2,354,1,2,1,1,40,122,226]
a: [3,380,1,2,1,354]
s: 380
i: [1,144,0,8,3,0,0,1] dt_k5_ordinal3
s: 144
i: [1,145,0,8,3,0,0,1] dt_k10_ordinal2
s: 145
d: [2,316,2,2,1,1,40,121,270]
a: [3,338,2,2,1,316]
s: 338
i: [1,137,0,9,3,0,0,255] t5_subset
s: 137
d: [2,319,2,2,1,1,40,121,275]
a: [3,348,2,2,1,319]
s: 348
i: [1,136,0,10,3,0,0,255] t4_subset
s: 136
i: [1,146,0,10,4,0,0,1] connectedness_r1_ordinal1
s: 146
d: [2,438,1,15,3,1,40,146,166]
d: [2,441,1,13,2,2,21,438,176]
a: [3,445,1,6,1,441]
s: 445
i: [1,147,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 147
d: [2,451,2,6,2,1,40,445,145]
d: [2,453,2,4,1,2,21,451,171]
a: [3,457,2,4,1,453]
s: 457
d: [2,472,3,0,0,3,52,176,171,457,144]
e: 472
d: [2,473,3,4,2,1,40,457,144]
d: [2,475,3,2,1,2,21,473,171]
d: [2,476,3,0,0,3,21,475,176]
e: 476
r: 445
r: 457
i: [1,148,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 148
a: [3,199,0,12,4,159]
s: 199
a: [3,456,2,4,1,453]
s: 456
i: [1,160,0,16,5,0,0,1] d5_ordinal3
s: 160
d: [2,452,2,6,2,1,40,445,141]
a: [3,466,2,4,1,452]
s: 466
i: [1,162,0,18,5,0,0,1] d5_ordinal3
s: 162
d: [2,512,0,14,4,0,21,162,144]
s: 512
d: [2,480,1,12,4,2,44,166,199]
d: [2,483,1,10,3,3,21,480,176]
d: [2,484,1,8,2,4,21,483,171]
a: [3,488,1,3,1,484]
s: 488
r: 466
r: 488
i: [1,161,0,20,5,0,0,1] d5_ordinal3
s: 161
a: [3,487,1,3,1,484]
s: 487
r: 487
d: [2,536,0,16,4,0,21,161,145]
s: 536
s: 488
d: [2,514,1,10,4,1,44,166,512]
d: [2,520,1,8,3,2,21,514,171]
d: [2,521,1,6,2,3,21,520,176]
a: [3,525,1,3,1,521]
s: 525
d: [2,559,2,0,0,3,52,176,525,206]
e: 559
d: [2,560,2,0,0,2,52,176,176,139,525,148]
e: 560
d: [2,561,2,0,0,2,52,176,176,525,525,146]
e: 561
d: [2,562,2,0,0,2,52,176,176,525,525,146]
e: 562
d: [2,563,2,2,1,2,40,525,206]
d: [2,570,2,0,0,3,21,563,176]
e: 570
r: 488
r: 525
s: 487
a: [3,524,1,3,1,521]
s: 524
a: [3,492,1,5,1,484]
s: 492
d: [2,572,2,7,3,1,40,487,146]
d: [2,573,2,5,2,2,21,572,171]
d: [2,574,2,3,1,3,21,573,176]
d: [2,580,2,5,2,1,44,492,120]
d: [2,591,2,3,1,2,21,580,171]
f: 591,574,562,521,484,303,159,112,108,109,110
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t62_ordinal3
% SZS output start Proof for ordinal3__t62_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => r1_ordinal1(X0,k10_ordinal2(X1,k5_ordinal3(X0,X1))))) [input t62_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => r1_ordinal1(X0,k10_ordinal2(X1,k5_ordinal3(X0,X1))))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => v3_ordinal1(X0)) [input cc3_ordinal1]
14. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1))) [input connectedness_r1_ordinal1]
16. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => ((~r1_ordinal1(X1,X0) => (k5_ordinal3(X0,X1) = X2 <=> k1_xboole_0 = X2)) & (r1_ordinal1(X1,X0) => (k5_ordinal3(X0,X1) = X2 <=> k10_ordinal2(X1,X2) = X0)))))) [input d5_ordinal3]
20. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k5_ordinal3(X0,X1))) [input dt_k5_ordinal3]
24. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
36. ! [X0] : (v3_ordinal1(X0) => k10_ordinal2(X0,k1_xboole_0) = X0) [input t27_ordinal2]
53. ? [X0] : (? [X1] : (~r1_ordinal1(X0,k10_ordinal2(X1,k5_ordinal3(X0,X1))) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
58. ! [X0] : (v3_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 8]
60. ! [X0] : (k10_ordinal2(X0,k1_xboole_0) = X0 | ~v3_ordinal1(X0)) [ennf transformation 36]
62. ! [X0] : (! [X1] : (! [X2] : ((((k5_ordinal3(X0,X1) = X2 <=> k1_xboole_0 = X2) | r1_ordinal1(X1,X0)) & ((k5_ordinal3(X0,X1) = X2 <=> k10_ordinal2(X1,X2) = X0) | ~r1_ordinal1(X1,X0))) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 16]
80. ! [X0,X1] : (v3_ordinal1(k5_ordinal3(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 20]
81. ! [X0,X1] : (v3_ordinal1(k5_ordinal3(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 80]
84. ! [X0,X1] : ((r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 14]
85. ! [X0,X1] : (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 84]
88. ? [X0] : (? [X1] : (~r1_ordinal1(X0,k10_ordinal2(X1,k5_ordinal3(X0,X1))) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (~r1_ordinal1(sK0,k10_ordinal2(X1,k5_ordinal3(sK0,X1))) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
89. ? [X1] : (~r1_ordinal1(sK0,k10_ordinal2(X1,k5_ordinal3(sK0,X1))) & v3_ordinal1(X1)) => (~r1_ordinal1(sK0,k10_ordinal2(sK1,k5_ordinal3(sK0,sK1))) & v3_ordinal1(sK1)) [choice axiom]
90. (~r1_ordinal1(sK0,k10_ordinal2(sK1,k5_ordinal3(sK0,sK1))) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 53,89,88]
91. ! [X0] : (! [X1] : (! [X2] : (((((k5_ordinal3(X0,X1) = X2 | k1_xboole_0 != X2) & (k1_xboole_0 = X2 | k5_ordinal3(X0,X1) != X2)) | r1_ordinal1(X1,X0)) & (((k5_ordinal3(X0,X1) = X2 | k10_ordinal2(X1,X2) != X0) & (k10_ordinal2(X1,X2) = X0 | k5_ordinal3(X0,X1) != X2)) | ~r1_ordinal1(X1,X0))) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [nnf transformation 62]
108. v3_ordinal1(sK0) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
109. v3_ordinal1(sK1) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
110. ~r1_ordinal1(sK0,k10_ordinal2(sK1,k5_ordinal3(sK0,sK1))) [cnf transformation 90] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
112. v1_xboole_0(k1_xboole_0) [cnf transformation 24] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
118. ~v1_xboole_0(X0) | v3_ordinal1(X0) [cnf transformation 58] {a:0,w:4,wCS:1600,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
120. k10_ordinal2(X0,k1_xboole_0) = X0 | ~v3_ordinal1(X0) [cnf transformation 60] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
123. k10_ordinal2(X1,X2) = X0 | k5_ordinal3(X0,X1) != X2 | ~r1_ordinal1(X1,X0) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 91] {a:0,w:19,wCS:7600,sine:1,thAx:0,allAx:1,thDist:-1}
126. k5_ordinal3(X0,X1) = X2 | k1_xboole_0 != X2 | r1_ordinal1(X1,X0) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 91] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
144. v3_ordinal1(k5_ordinal3(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 81] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
146. r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 85] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
159. k5_ordinal3(X0,X1) = k1_xboole_0 | r1_ordinal1(X1,X0) | ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [equality resolution 126] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:1,thDist:-1}
162. k10_ordinal2(X1,k5_ordinal3(X0,X1)) = X0 | ~r1_ordinal1(X1,X0) | ~v3_ordinal1(k5_ordinal3(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [equality resolution 123] {a:0,w:18,wCS:7200,sine:1,thAx:0,allAx:1,thDist:-1}
164. 1 <=> r1_ordinal1(sK0,k10_ordinal2(sK1,k5_ordinal3(sK0,sK1))) [avatar definition]
166. ~r1_ordinal1(sK0,k10_ordinal2(sK1,k5_ordinal3(sK0,sK1))) <- (~1) [avatar component clause 164] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
167. ~1 [avatar split clause 110,164]
169. 2 <=> v3_ordinal1(sK1) [avatar definition]
171. v3_ordinal1(sK1) <- (2) [avatar component clause 169] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
172. 2 [avatar split clause 109,169]
174. 3 <=> v3_ordinal1(sK0) [avatar definition]
176. v3_ordinal1(sK0) <- (3) [avatar component clause 174] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
177. 3 [avatar split clause 108,174]
184. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
186. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 184] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
187. 5 [avatar split clause 112,184]
194. 7 <=> v3_ordinal1(k1_xboole_0) [avatar definition]
198. 8 <=> ! [X1,X0] : (k5_ordinal3(X0,X1) = k1_xboole_0 | ~v3_ordinal1(X0) | ~v3_ordinal1(X1) | r1_ordinal1(X1,X0)) [avatar definition]
199. k5_ordinal3(X0,X1) = k1_xboole_0 | ~v3_ordinal1(X0) | ~v3_ordinal1(X1) | r1_ordinal1(X1,X0) <- (8) [avatar component clause 198] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
200. ~7 | 8 [avatar split clause 159,198,194]
303. v3_ordinal1(k1_xboole_0) <- (5) [resolution 118,186] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
307. 7 | ~5 [avatar split clause 303,184,194]
480. ~r1_ordinal1(sK0,k10_ordinal2(sK1,k1_xboole_0)) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK1) | r1_ordinal1(sK1,sK0) <- (~1, 8) [superposition 166,199] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
483. ~r1_ordinal1(sK0,k10_ordinal2(sK1,k1_xboole_0)) | ~v3_ordinal1(sK1) | r1_ordinal1(sK1,sK0) <- (~1, 3, 8) [subsumption resolution 480,176] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
484. ~r1_ordinal1(sK0,k10_ordinal2(sK1,k1_xboole_0)) | r1_ordinal1(sK1,sK0) <- (~1, 2, 3, 8) [subsumption resolution 483,171] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
486. 47 <=> r1_ordinal1(sK1,sK0) [avatar definition]
487. ~r1_ordinal1(sK1,sK0) <- (~47) [avatar component clause 486] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
490. 48 <=> r1_ordinal1(sK0,k10_ordinal2(sK1,k1_xboole_0)) [avatar definition]
492. ~r1_ordinal1(sK0,k10_ordinal2(sK1,k1_xboole_0)) <- (~48) [avatar component clause 490] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
493. 47 | ~48 | 1 | ~2 | ~3 | ~8 [avatar split clause 484,198,174,169,164,490,486]
512. k10_ordinal2(X1,k5_ordinal3(X0,X1)) = X0 | ~r1_ordinal1(X1,X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [subsumption resolution 162,144] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
514. ~r1_ordinal1(sK0,sK0) | ~r1_ordinal1(sK1,sK0) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~1) [superposition 166,512] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
520. ~r1_ordinal1(sK0,sK0) | ~r1_ordinal1(sK1,sK0) | ~v3_ordinal1(sK0) <- (~1, 2) [subsumption resolution 514,171] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
521. ~r1_ordinal1(sK0,sK0) | ~r1_ordinal1(sK1,sK0) <- (~1, 2, 3) [subsumption resolution 520,176] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
523. 51 <=> r1_ordinal1(sK0,sK0) [avatar definition]
525. ~r1_ordinal1(sK0,sK0) <- (~51) [avatar component clause 523] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
526. ~47 | ~51 | 1 | ~2 | ~3 [avatar split clause 521,174,169,164,523,486]
541. 53 <=> r1_ordinal1(sK0,sK1) [avatar definition]
562. $false <- (3, ~51) [unit resulting resolution 176,176,525,525,146] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
569. ~3 | 51 [avatar contradiction clause 562]
572. r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~47) [resolution 487,146] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
573. r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK0) <- (2, ~47) [subsumption resolution 572,171] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
574. r1_ordinal1(sK0,sK1) <- (2, 3, ~47) [subsumption resolution 573,176] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
575. 53 | ~2 | ~3 | 47 [avatar split clause 574,486,174,169,541]
580. ~r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK1) <- (~48) [superposition 492,120] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
591. ~r1_ordinal1(sK0,sK1) <- (2, ~48) [subsumption resolution 580,171] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
592. ~53 | ~2 | 48 [avatar split clause 591,490,169,541]
593. $false [avatar sat refutation 167,172,177,187,200,307,493,526,569,575,592] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:389,thDist:-389}
% SZS output end Proof for ordinal3__t62_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 42
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 51
% Generated clauses: 260
% Active clauses: 81
% Passive clauses: 121
% Final active clauses: 70
% Final passive clauses: 36
% Main loop iterations started: 85
% 
% >>> Simplifying Inferences
% Duplicate literals: 18
% Fw subsumption resolutions: 28
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Forward subsumptions: 22
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 69
% Unit resulting resolution: 7
% Factoring: 1
% Forward superposition: 15
% Backward superposition: 10
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 14
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 69
% SAT solver unit clauses: 16
% SAT solver binary clauses: 35
% 
% Memory used [KB]: 5245
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

