i: [1,45,0,5,1,1,0,0] t126_zfmisc_1
a: [3,62,0,5,1,45]
s: 62
i: [1,46,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,67,0,2,1,46]
s: 67
i: [1,47,0,2,1,0,0,255] fc1_xboole_0
a: [3,72,0,2,1,47]
s: 72
i: [1,56,0,2,1,0,0,2] rc2_xboole_0
a: [3,82,0,2,1,56]
s: 82
i: [1,57,0,2,1,0,0,2] rc1_xboole_0
a: [3,87,0,2,1,57]
s: 87
i: [1,58,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,92,0,2,1,58]
s: 92
i: [1,44,0,9,3,1,0,0] t126_zfmisc_1
s: 44
i: [1,48,0,3,1,0,0,255] d2_xboole_0
a: [3,77,0,3,1,48]
s: 77
r: 67
i: [1,49,0,5,2,0,0,255] t6_boole
s: 49
i: [1,53,0,5,2,0,0,1] t7_boole
s: 53
i: [1,54,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 54
i: [1,55,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 55
i: [1,52,0,7,3,0,0,255] t8_boole
s: 52
d: [2,95,1,7,3,0,44,49,49]
s: 95
i: [1,50,0,9,2,0,0,1] t80_zfmisc_1
s: 50
i: [1,51,0,9,2,0,0,1] t80_zfmisc_1
s: 51
d: [2,97,1,12,3,0,40,50,44]
s: 97
d: [2,99,1,6,2,0,40,50,53]
s: 99
d: [2,102,2,2,1,1,40,99,62]
a: [3,107,2,2,1,102]
s: 107
d: [2,98,1,9,2,0,40,50,54]
s: 98
d: [2,100,2,12,3,0,40,97,55]
s: 100
d: [2,109,3,11,3,0,40,100,51]
d: [2,111,3,7,2,0,18,109]
s: 111
d: [2,101,2,17,3,0,40,97,51]
s: 101
d: [2,112,4,7,2,0,40,111,55]
s: 112
d: [2,103,2,6,2,0,40,99,55]
s: 103
d: [2,123,3,2,1,1,40,103,62]
a: [3,129,3,2,1,123]
s: 129
d: [2,115,5,4,1,1,40,112,62]
a: [3,121,5,4,1,115]
s: 121
d: [2,113,3,16,3,0,40,101,50]
d: [2,114,3,10,2,0,18,113]
s: 114
d: [2,131,4,10,2,0,40,114,55]
s: 131
d: [2,132,5,10,2,0,40,131,51]
d: [2,134,5,5,1,0,18,132]
a: [3,138,5,5,1,134]
s: 138
d: [2,140,6,0,0,2,52,62,138,55]
e: 140
d: [2,141,6,5,1,1,40,138,55]
d: [2,143,6,0,0,2,21,141,62]
e: 143
f: 143,134,45
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t126_zfmisc_1
% SZS output start Proof for zfmisc_1__t126_zfmisc_1
1. ! [X0] : ! [X1] : (! [X2] : ! [X3] : ((r2_hidden(X3,X1) & r2_hidden(X2,X0)) => r1_xboole_0(X2,X3)) => r1_xboole_0(k3_tarski(X0),k3_tarski(X1))) [input t126_zfmisc_1]
2. ~! [X0] : ! [X1] : (! [X2] : ! [X3] : ((r2_hidden(X3,X1) & r2_hidden(X2,X0)) => r1_xboole_0(X2,X3)) => r1_xboole_0(k3_tarski(X0),k3_tarski(X1))) [negated conjecture 1]
14. ! [X0,X1] : (r1_xboole_0(X0,X1) => r1_xboole_0(X1,X0)) [input symmetry_r1_xboole_0]
17. ! [X0] : ! [X1] : (! [X2] : (r2_hidden(X2,X0) => r1_xboole_0(X2,X1)) => r1_xboole_0(k3_tarski(X0),X1)) [input t80_zfmisc_1]
19. ~! [X0,X1] : (! [X2,X3] : ((r2_hidden(X3,X1) & r2_hidden(X2,X0)) => r1_xboole_0(X2,X3)) => r1_xboole_0(k3_tarski(X0),k3_tarski(X1))) [flattening 2]
20. ! [X0,X1] : (! [X2] : (r2_hidden(X2,X0) => r1_xboole_0(X2,X1)) => r1_xboole_0(k3_tarski(X0),X1)) [flattening 17]
26. ? [X0,X1] : (~r1_xboole_0(k3_tarski(X0),k3_tarski(X1)) & ! [X2,X3] : (r1_xboole_0(X2,X3) | (~r2_hidden(X3,X1) | ~r2_hidden(X2,X0)))) [ennf transformation 19]
27. ? [X0,X1] : (~r1_xboole_0(k3_tarski(X0),k3_tarski(X1)) & ! [X2,X3] : (r1_xboole_0(X2,X3) | ~r2_hidden(X3,X1) | ~r2_hidden(X2,X0))) [flattening 26]
29. ! [X0,X1] : (r1_xboole_0(k3_tarski(X0),X1) | ? [X2] : (~r1_xboole_0(X2,X1) & r2_hidden(X2,X0))) [ennf transformation 20]
33. ! [X0,X1] : (r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 14]
34. ? [X0,X1] : (~r1_xboole_0(k3_tarski(X0),k3_tarski(X1)) & ! [X2,X3] : (r1_xboole_0(X2,X3) | ~r2_hidden(X3,X1) | ~r2_hidden(X2,X0))) => (~r1_xboole_0(k3_tarski(sK0),k3_tarski(sK1)) & ! [X3,X2] : (r1_xboole_0(X2,X3) | ~r2_hidden(X3,sK1) | ~r2_hidden(X2,sK0))) [choice axiom]
35. ~r1_xboole_0(k3_tarski(sK0),k3_tarski(sK1)) & ! [X2,X3] : (r1_xboole_0(X2,X3) | ~r2_hidden(X3,sK1) | ~r2_hidden(X2,sK0)) [skolemisation 27,34]
36. ! [X1,X0] : (? [X2] : (~r1_xboole_0(X2,X1) & r2_hidden(X2,X0)) => (~r1_xboole_0(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) [choice axiom]
37. ! [X0,X1] : (r1_xboole_0(k3_tarski(X0),X1) | (~r1_xboole_0(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) [skolemisation 29,36]
44. ~r2_hidden(X3,sK1) | r1_xboole_0(X2,X3) | ~r2_hidden(X2,sK0) [cnf transformation 35] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. ~r1_xboole_0(k3_tarski(sK0),k3_tarski(sK1)) [cnf transformation 35] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. r2_hidden(sK2(X0,X1),X0) | r1_xboole_0(k3_tarski(X0),X1) [cnf transformation 37] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
51. ~r1_xboole_0(sK2(X0,X1),X1) | r1_xboole_0(k3_tarski(X0),X1) [cnf transformation 37] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
55. ~r1_xboole_0(X0,X1) | r1_xboole_0(X1,X0) [cnf transformation 33] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. 1 <=> r1_xboole_0(k3_tarski(sK0),k3_tarski(sK1)) [avatar definition]
62. ~r1_xboole_0(k3_tarski(sK0),k3_tarski(sK1)) <- (~1) [avatar component clause 60] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. ~1 [avatar split clause 45,60]
97. r1_xboole_0(X1,sK2(sK1,X0)) | r1_xboole_0(k3_tarski(sK1),X0) | ~r2_hidden(X1,sK0) [resolution 50,44] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
101. ~r2_hidden(sK2(X3,sK2(sK1,X2)),sK0) | r1_xboole_0(k3_tarski(sK1),X2) | r1_xboole_0(k3_tarski(X3),sK2(sK1,X2)) [resolution 97,51] {a:2,w:17,wCS:1700,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
113. r1_xboole_0(k3_tarski(sK1),X0) | r1_xboole_0(k3_tarski(sK0),sK2(sK1,X0)) | r1_xboole_0(k3_tarski(sK0),sK2(sK1,X0)) [resolution 101,50] {a:3,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
114. r1_xboole_0(k3_tarski(sK0),sK2(sK1,X0)) | r1_xboole_0(k3_tarski(sK1),X0) [duplicate literal removal 113] {a:3,w:10,wCS:1000,nSel:1,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
131. r1_xboole_0(sK2(sK1,X0),k3_tarski(sK0)) | r1_xboole_0(k3_tarski(sK1),X0) [resolution 114,55] {a:4,w:10,wCS:1000,nSel:1,goal:1,sine:0,thAx:0,allAx:9,thDist:-9}
132. r1_xboole_0(k3_tarski(sK1),k3_tarski(sK0)) | r1_xboole_0(k3_tarski(sK1),k3_tarski(sK0)) [resolution 131,51] {a:5,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:11,thDist:-11}
134. r1_xboole_0(k3_tarski(sK1),k3_tarski(sK0)) [duplicate literal removal 132] {a:5,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:11,thDist:-11}
136. 11 <=> r1_xboole_0(k3_tarski(sK1),k3_tarski(sK0)) [avatar definition]
138. r1_xboole_0(k3_tarski(sK1),k3_tarski(sK0)) <- (11) [avatar component clause 136] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
139. 11 [avatar split clause 134,136]
141. r1_xboole_0(k3_tarski(sK0),k3_tarski(sK1)) <- (11) [resolution 138,55] {a:6,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
143. $false <- (~1, 11) [subsumption resolution 141,62] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
144. 1 | ~11 [avatar contradiction clause 143]
145. $false [avatar sat refutation 63,139,144] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:83,thDist:-83}
% SZS output end Proof for zfmisc_1__t126_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 55
% Active clauses: 29
% Passive clauses: 31
% Final active clauses: 28
% Final passive clauses: 1
% Main loop iterations started: 30
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 21
% Unit resulting resolution: 1
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 8
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% passive container maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

