i: [1,117,0,6,2,1,0,0] t6_interva1
a: [3,171,0,3,1,117]
s: 171
i: [1,114,0,4,1,1,0,0] t6_interva1
a: [3,190,0,4,1,114]
s: 190
i: [1,113,0,4,1,1,0,0] t6_interva1
a: [3,195,0,4,1,113]
s: 195
i: [1,112,0,4,1,1,0,0] t6_interva1
a: [3,200,0,4,1,112]
s: 200
i: [1,111,0,4,1,1,0,0] t6_interva1
a: [3,205,0,4,1,111]
s: 205
i: [1,116,0,6,1,1,0,0] t6_interva1
a: [3,180,0,6,1,116]
s: 180
i: [1,118,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,210,0,2,1,118]
s: 210
i: [1,119,0,2,1,0,0,1] fc1_xboole_0
a: [3,215,0,2,1,119]
s: 215
i: [1,163,0,2,1,0,0,2] rc2_xboole_0
a: [3,225,0,2,1,163]
s: 225
i: [1,164,0,2,1,0,0,2] rc1_xboole_0
a: [3,230,0,2,1,164]
s: 230
i: [1,115,0,9,1,1,0,0] t6_interva1
a: [3,185,0,9,1,115]
s: 185
i: [1,121,0,3,1,0,0,2] fc1_subset_1
s: 121
i: [1,152,0,3,1,0,0,2] rc2_subset_1
s: 152
i: [1,155,0,3,1,0,0,2] reflexivity_r1_tarski
s: 155
i: [1,165,0,3,1,0,0,2] d10_xboole_0
s: 165
i: [1,166,0,3,1,0,0,2] d10_xboole_0
s: 166
i: [1,120,0,3,1,0,0,255] d2_xboole_0
a: [3,220,0,3,1,120]
s: 220
r: 210
i: [1,150,0,4,1,0,0,1] existence_m1_subset_1
s: 150
i: [1,154,0,4,1,0,0,255] rc3_subset_1
s: 154
i: [1,125,0,5,2,0,0,2] rc1_subset_1
s: 125
i: [1,128,0,5,2,0,0,1] t6_boole
s: 128
i: [1,145,0,5,2,0,0,2] t7_boole
s: 145
i: [1,151,0,5,1,0,0,2] rc2_subset_1
s: 151
i: [1,153,0,5,1,0,0,255] rc3_subset_1
s: 153
i: [1,127,0,6,2,0,0,255] rc4_subset_1
s: 127
i: [1,131,0,6,2,0,0,2] t1_subset
s: 131
i: [1,156,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 156
i: [1,124,0,7,2,0,0,2] rc1_subset_1
s: 124
i: [1,126,0,7,2,0,0,255] rc4_subset_1
s: 126
i: [1,142,0,7,2,0,0,2] t3_subset
s: 142
i: [1,143,0,7,2,0,0,2] t3_subset
s: 143
d: [2,235,1,3,1,1,40,142,205]
a: [3,247,1,3,1,235]
s: 247
i: [1,144,0,7,3,0,0,1] t8_boole
s: 144
d: [2,236,1,3,1,1,40,142,200]
a: [3,252,1,3,1,236]
s: 252
i: [1,146,0,7,2,0,0,2] t2_tarski
s: 146
i: [1,129,0,8,3,0,0,2] cc1_subset_1
s: 129
d: [2,268,1,4,2,1,40,129,205]
a: [3,281,1,2,1,268]
s: 281
i: [1,132,0,8,3,0,0,2] t2_subset
s: 132
d: [2,237,1,3,1,1,40,142,195]
a: [3,257,1,3,1,237]
s: 257
i: [1,130,0,9,3,0,0,255] cc4_subset_1
s: 130
i: [1,141,0,9,3,0,0,2] d10_xboole_0
s: 141
d: [2,238,1,3,1,1,40,142,190]
a: [3,262,1,3,1,238]
s: 262
i: [1,149,0,9,3,0,0,2] t5_subset
s: 149
d: [2,352,2,6,2,1,40,141,247]
a: [3,358,2,3,1,352]
s: 358
i: [1,147,0,10,3,0,0,2] t2_tarski
s: 147
i: [1,148,0,10,3,0,0,2] t4_subset
s: 148
d: [2,353,2,6,2,1,40,141,252]
a: [3,367,2,3,1,353]
s: 367
i: [1,122,0,11,4,0,0,255] cc3_subset_1
s: 122
a: [3,284,1,2,1,268]
s: 284
i: [1,123,0,11,4,0,0,255] cc2_subset_1
s: 123
i: [1,157,0,16,3,0,0,1] dt_k1_interva1
s: 157
d: [2,269,1,4,2,1,40,129,200]
a: [3,289,1,2,1,269]
s: 289
i: [1,133,0,17,3,0,0,255] d1_interva1
s: 133
d: [2,270,1,4,2,1,40,129,195]
a: [3,294,1,2,1,270]
s: 294
i: [1,136,0,17,4,0,0,1] t1_interva1
s: 136
i: [1,137,0,17,4,0,0,1] t1_interva1
s: 137
d: [2,271,1,4,2,1,40,129,190]
a: [3,299,1,2,1,271]
s: 299
i: [1,134,0,20,4,0,0,1] t2_interva1
s: 134
d: [2,354,2,6,2,1,40,141,257]
a: [3,376,2,3,1,354]
s: 376
i: [1,135,0,20,4,0,0,1] t2_interva1
s: 135
i: [1,138,0,20,5,0,0,1] t1_interva1
s: 138
d: [2,382,2,6,2,1,40,262,141]
a: [3,386,2,3,1,382]
s: 386
i: [1,159,0,21,4,0,0,255] fraenkel_a_3_0_interva1
s: 159
d: [2,302,1,7,2,1,40,132,205]
d: [2,312,1,4,1,1,21,302,121]
a: [3,316,1,4,1,312]
s: 316
i: [1,160,0,21,4,0,0,255] fraenkel_a_3_0_interva1
s: 160
i: [1,161,0,21,4,0,0,255] fraenkel_a_3_0_interva1
s: 161
d: [2,303,1,7,2,1,40,132,200]
d: [2,318,1,4,1,1,21,303,121]
a: [3,322,1,4,1,318]
s: 322
i: [1,158,0,22,4,0,0,255] fraenkel_a_3_0_interva1
s: 158
d: [2,304,1,7,2,1,40,132,195]
d: [2,324,1,4,1,1,21,304,121]
a: [3,328,1,4,1,324]
s: 328
i: [1,167,0,24,6,0,0,255] fraenkel_a_3_0_interva1
s: 167
d: [2,305,1,7,2,1,40,132,190]
d: [2,330,1,4,1,1,21,305,121]
a: [3,334,1,4,1,330]
s: 334
d: [2,518,2,4,1,1,40,316,156]
a: [3,524,2,4,1,518]
s: 524
d: [2,403,1,6,2,1,40,148,205]
s: 403
d: [2,571,2,7,2,1,40,403,146]
a: [3,574,2,3,1,571]
s: 574
r: 524
r: 403
d: [2,580,3,0,0,2,52,150,284,574,132]
e: 580
r: 574
s: 403
d: [2,404,1,6,2,1,40,148,200]
s: 404
d: [2,585,2,7,2,1,40,404,146]
a: [3,588,2,3,1,585]
s: 588
r: 404
d: [2,594,3,0,0,2,52,150,289,588,132]
e: 594
r: 289
d: [2,414,1,7,3,1,40,122,200]
d: [2,431,1,5,2,2,21,414,281]
a: [3,435,1,3,1,431]
s: 435
s: 524
d: [2,405,1,6,2,1,40,148,195]
s: 405
d: [2,406,1,6,2,1,40,148,190]
s: 406
d: [2,597,2,7,2,1,40,405,146]
a: [3,600,2,3,1,597]
s: 600
r: 405
d: [2,615,3,0,0,2,52,150,294,600,132]
e: 615
r: 294
d: [2,415,1,7,3,1,40,122,195]
d: [2,437,1,5,2,2,21,415,281]
a: [3,441,1,3,1,437]
s: 441
d: [2,606,2,7,2,1,40,406,146]
a: [3,609,2,3,1,606]
s: 609
r: 406
d: [2,618,3,0,0,2,52,150,299,609,132]
e: 618
r: 299
d: [2,416,1,7,3,1,40,122,190]
d: [2,443,1,5,2,2,21,416,281]
a: [3,447,1,3,1,443]
s: 447
d: [2,481,1,20,4,1,44,134,185]
d: [2,484,1,14,3,2,21,481,180]
d: [2,485,1,10,2,3,21,484,190]
d: [2,486,1,6,1,4,21,485,195]
a: [3,490,1,6,1,486]
s: 490
d: [2,621,2,11,3,1,40,490,137]
d: [2,626,2,7,2,2,21,621,200]
d: [2,627,2,3,1,3,21,626,205]
a: [3,631,2,3,1,627]
s: 631
d: [2,497,1,20,4,1,44,135,185]
d: [2,500,1,14,3,2,21,497,180]
d: [2,501,1,10,2,3,21,500,190]
d: [2,502,1,6,1,4,21,501,195]
a: [3,506,1,6,1,502]
s: 506
d: [2,622,2,11,3,1,40,490,136]
d: [2,633,2,7,2,2,21,622,200]
d: [2,634,2,3,1,3,21,633,205]
a: [3,638,2,3,1,634]
s: 638
d: [2,458,1,16,3,1,44,157,185]
d: [2,461,1,12,2,2,21,458,190]
d: [2,462,1,8,1,3,21,461,195]
a: [3,466,1,8,1,462]
s: 466
d: [2,469,1,17,4,1,44,136,185]
d: [2,470,1,13,3,2,21,469,190]
d: [2,471,1,9,2,3,21,470,195]
s: 471
r: 171
a: [3,170,0,3,1,117]
s: 170
r: 185
r: 195
r: 257
r: 328
r: 376
r: 441
r: 471
r: 490
r: 600
r: 631
r: 638
r: 284
a: [3,175,0,3,1,117]
s: 175
d: [2,413,1,7,3,1,40,122,205]
d: [2,425,1,5,2,2,21,413,281]
a: [3,429,1,3,1,425]
s: 429
d: [2,742,0,9,1,2,23,185,170]
a: [3,756,0,9,1,742]
s: 756
d: [2,749,1,6,1,2,23,490,170]
a: [3,762,1,6,1,749]
s: 762
s: 574
r: 524
r: 403
d: [2,473,1,17,4,1,44,137,185]
d: [2,474,1,13,3,2,21,473,190]
d: [2,475,1,9,2,3,21,474,195]
s: 475
d: [2,660,2,11,3,1,40,506,137]
d: [2,665,2,7,2,2,21,660,200]
d: [2,666,2,3,1,3,21,665,205]
a: [3,670,2,3,1,666]
s: 670
d: [2,716,2,17,4,3,40,471,134]
d: [2,722,2,11,3,4,21,716,180]
d: [2,723,2,7,2,5,21,722,200]
d: [2,724,2,3,1,6,21,723,205]
d: [2,685,3,6,2,1,40,638,141]
d: [2,807,2,17,4,3,40,475,135]
d: [2,816,2,11,3,4,21,807,180]
d: [2,817,2,7,2,5,21,816,200]
d: [2,818,2,3,1,6,21,817,205]
d: [2,832,3,6,2,1,40,670,141]
d: [2,833,3,3,1,2,21,832,175]
f: 833,818,685,724,666,634,502,486,111,112,113,114,115,116,117
% Refutation found. Thanks to Tanya!
% SZS status Theorem for interva1__t6_interva1
% SZS output start Proof for interva1__t6_interva1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(X0)) => ! [X4] : (m1_subset_1(X4,k1_zfmisc_1(X0)) => (k1_interva1(X0,X1,X2) = k1_interva1(X0,X3,X4) => ((X2 = X4 & X1 = X3) | k1_interva1(X0,X1,X2) = k1_xboole_0)))))) [input t6_interva1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(X0)) => ! [X4] : (m1_subset_1(X4,k1_zfmisc_1(X0)) => (k1_interva1(X0,X1,X2) = k1_interva1(X0,X3,X4) => ((X2 = X4 & X1 = X3) | k1_interva1(X0,X1,X2) = k1_xboole_0)))))) [negated conjecture 1]
8. ! [X0] : ! [X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [input d10_xboole_0]
27. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => ! [X3] : (r2_hidden(X3,k1_interva1(X0,X1,X2)) <=> (r1_tarski(X3,X2) & r1_tarski(X1,X3))))) [input t1_interva1]
29. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => (k1_interva1(X0,X1,X2) != k1_xboole_0 => (r2_hidden(X2,k1_interva1(X0,X1,X2)) & r2_hidden(X1,k1_interva1(X0,X1,X2)))))) [input t2_interva1]
38. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(X0)) => ! [X4] : (m1_subset_1(X4,k1_zfmisc_1(X0)) => (k1_interva1(X0,X1,X2) = k1_interva1(X0,X3,X4) => ((X2 = X4 & X1 = X3) | k1_interva1(X0,X1,X2) = k1_xboole_0)))))) [flattening 2]
42. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => (k1_interva1(X0,X1,X2) != k1_xboole_0 => (r2_hidden(X2,k1_interva1(X0,X1,X2)) & r2_hidden(X1,k1_interva1(X0,X1,X2)))))) [flattening 29]
43. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => ! [X3] : (r2_hidden(X3,k1_interva1(X0,X1,X2)) <=> (r1_tarski(X3,X2) & r1_tarski(X1,X3))))) [flattening 27]
44. ! [X0,X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [flattening 8]
52. ? [X0,X1] : (? [X2] : (? [X3] : (? [X4] : ((((X2 != X4 | X1 != X3) & k1_interva1(X0,X1,X2) != k1_xboole_0) & k1_interva1(X0,X1,X2) = k1_interva1(X0,X3,X4)) & m1_subset_1(X4,k1_zfmisc_1(X0))) & m1_subset_1(X3,k1_zfmisc_1(X0))) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 38]
53. ? [X0,X1] : (? [X2] : (? [X3] : (? [X4] : ((X2 != X4 | X1 != X3) & k1_interva1(X0,X1,X2) != k1_xboole_0 & k1_interva1(X0,X1,X2) = k1_interva1(X0,X3,X4) & m1_subset_1(X4,k1_zfmisc_1(X0))) & m1_subset_1(X3,k1_zfmisc_1(X0))) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 52]
67. ! [X0,X1] : (! [X2] : (((r2_hidden(X2,k1_interva1(X0,X1,X2)) & r2_hidden(X1,k1_interva1(X0,X1,X2))) | k1_interva1(X0,X1,X2) = k1_xboole_0) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 42]
68. ! [X0,X1] : (! [X2] : ((r2_hidden(X2,k1_interva1(X0,X1,X2)) & r2_hidden(X1,k1_interva1(X0,X1,X2))) | k1_interva1(X0,X1,X2) = k1_xboole_0 | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 67]
69. ! [X0,X1] : (! [X2] : (! [X3] : (r2_hidden(X3,k1_interva1(X0,X1,X2)) <=> (r1_tarski(X3,X2) & r1_tarski(X1,X3))) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 43]
81. ? [X0,X1] : (? [X2] : (? [X3] : (? [X4] : ((X2 != X4 | X1 != X3) & k1_interva1(X0,X1,X2) != k1_xboole_0 & k1_interva1(X0,X1,X2) = k1_interva1(X0,X3,X4) & m1_subset_1(X4,k1_zfmisc_1(X0))) & m1_subset_1(X3,k1_zfmisc_1(X0))) & m1_subset_1(X2,k1_zfmisc_1(X0))) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (? [X2] : (? [X3] : (? [X4] : ((X2 != X4 | sK1 != X3) & k1_xboole_0 != k1_interva1(sK0,sK1,X2) & k1_interva1(sK0,sK1,X2) = k1_interva1(sK0,X3,X4) & m1_subset_1(X4,k1_zfmisc_1(sK0))) & m1_subset_1(X3,k1_zfmisc_1(sK0))) & m1_subset_1(X2,k1_zfmisc_1(sK0))) & m1_subset_1(sK1,k1_zfmisc_1(sK0))) [choice axiom]
82. ? [X2] : (? [X3] : (? [X4] : ((X2 != X4 | sK1 != X3) & k1_xboole_0 != k1_interva1(sK0,sK1,X2) & k1_interva1(sK0,sK1,X2) = k1_interva1(sK0,X3,X4) & m1_subset_1(X4,k1_zfmisc_1(sK0))) & m1_subset_1(X3,k1_zfmisc_1(sK0))) & m1_subset_1(X2,k1_zfmisc_1(sK0))) => (? [X3] : (? [X4] : ((sK2 != X4 | sK1 != X3) & k1_xboole_0 != k1_interva1(sK0,sK1,sK2) & k1_interva1(sK0,X3,X4) = k1_interva1(sK0,sK1,sK2) & m1_subset_1(X4,k1_zfmisc_1(sK0))) & m1_subset_1(X3,k1_zfmisc_1(sK0))) & m1_subset_1(sK2,k1_zfmisc_1(sK0))) [choice axiom]
83. ? [X3] : (? [X4] : ((sK2 != X4 | sK1 != X3) & k1_xboole_0 != k1_interva1(sK0,sK1,sK2) & k1_interva1(sK0,X3,X4) = k1_interva1(sK0,sK1,sK2) & m1_subset_1(X4,k1_zfmisc_1(sK0))) & m1_subset_1(X3,k1_zfmisc_1(sK0))) => (? [X4] : ((sK2 != X4 | sK1 != sK3) & k1_xboole_0 != k1_interva1(sK0,sK1,sK2) & k1_interva1(sK0,sK1,sK2) = k1_interva1(sK0,sK3,X4) & m1_subset_1(X4,k1_zfmisc_1(sK0))) & m1_subset_1(sK3,k1_zfmisc_1(sK0))) [choice axiom]
84. ? [X4] : ((sK2 != X4 | sK1 != sK3) & k1_xboole_0 != k1_interva1(sK0,sK1,sK2) & k1_interva1(sK0,sK1,sK2) = k1_interva1(sK0,sK3,X4) & m1_subset_1(X4,k1_zfmisc_1(sK0))) => ((sK2 != sK4 | sK1 != sK3) & k1_xboole_0 != k1_interva1(sK0,sK1,sK2) & k1_interva1(sK0,sK1,sK2) = k1_interva1(sK0,sK3,sK4) & m1_subset_1(sK4,k1_zfmisc_1(sK0))) [choice axiom]
85. ((((sK2 != sK4 | sK1 != sK3) & k1_xboole_0 != k1_interva1(sK0,sK1,sK2) & k1_interva1(sK0,sK1,sK2) = k1_interva1(sK0,sK3,sK4) & m1_subset_1(sK4,k1_zfmisc_1(sK0))) & m1_subset_1(sK3,k1_zfmisc_1(sK0))) & m1_subset_1(sK2,k1_zfmisc_1(sK0))) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) [skolemisation 53,84,83,82,81]
90. ! [X0,X1] : (! [X2] : (! [X3] : ((r2_hidden(X3,k1_interva1(X0,X1,X2)) | (~r1_tarski(X3,X2) | ~r1_tarski(X1,X3))) & ((r1_tarski(X3,X2) & r1_tarski(X1,X3)) | ~r2_hidden(X3,k1_interva1(X0,X1,X2)))) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [nnf transformation 69]
91. ! [X0,X1] : (! [X2] : (! [X3] : ((r2_hidden(X3,k1_interva1(X0,X1,X2)) | ~r1_tarski(X3,X2) | ~r1_tarski(X1,X3)) & ((r1_tarski(X3,X2) & r1_tarski(X1,X3)) | ~r2_hidden(X3,k1_interva1(X0,X1,X2)))) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) [flattening 90]
92. ! [X0,X1] : ((X0 = X1 | (~r1_tarski(X1,X0) | ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [nnf transformation 44]
93. ! [X0,X1] : ((X0 = X1 | ~r1_tarski(X1,X0) | ~r1_tarski(X0,X1)) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [flattening 92]
111. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 85] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
112. m1_subset_1(sK2,k1_zfmisc_1(sK0)) [cnf transformation 85] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
113. m1_subset_1(sK3,k1_zfmisc_1(sK0)) [cnf transformation 85] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
114. m1_subset_1(sK4,k1_zfmisc_1(sK0)) [cnf transformation 85] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
115. k1_interva1(sK0,sK1,sK2) = k1_interva1(sK0,sK3,sK4) [cnf transformation 85] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
116. k1_xboole_0 != k1_interva1(sK0,sK1,sK2) [cnf transformation 85] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
117. sK2 != sK4 | sK1 != sK3 [cnf transformation 85] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
134. r2_hidden(X1,k1_interva1(X0,X1,X2)) | k1_interva1(X0,X1,X2) = k1_xboole_0 | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 68] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
135. r2_hidden(X2,k1_interva1(X0,X1,X2)) | k1_interva1(X0,X1,X2) = k1_xboole_0 | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 68] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
136. ~r2_hidden(X3,k1_interva1(X0,X1,X2)) | r1_tarski(X1,X3) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 91] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
137. ~r2_hidden(X3,k1_interva1(X0,X1,X2)) | r1_tarski(X3,X2) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) [cnf transformation 91] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
141. ~r1_tarski(X1,X0) | X0 = X1 | ~r1_tarski(X0,X1) [cnf transformation 93] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
169. 1 <=> sK1 = sK3 [avatar definition]
173. 2 <=> sK2 = sK4 [avatar definition]
175. sK2 != sK4 <- (~2) [avatar component clause 173] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
176. ~1 | ~2 [avatar split clause 117,173,169]
178. 3 <=> k1_xboole_0 = k1_interva1(sK0,sK1,sK2) [avatar definition]
180. k1_xboole_0 != k1_interva1(sK0,sK1,sK2) <- (~3) [avatar component clause 178] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
181. ~3 [avatar split clause 116,178]
183. 4 <=> k1_interva1(sK0,sK1,sK2) = k1_interva1(sK0,sK3,sK4) [avatar definition]
185. k1_interva1(sK0,sK1,sK2) = k1_interva1(sK0,sK3,sK4) <- (4) [avatar component clause 183] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
186. 4 [avatar split clause 115,183]
188. 5 <=> m1_subset_1(sK4,k1_zfmisc_1(sK0)) [avatar definition]
190. m1_subset_1(sK4,k1_zfmisc_1(sK0)) <- (5) [avatar component clause 188] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
191. 5 [avatar split clause 114,188]
193. 6 <=> m1_subset_1(sK3,k1_zfmisc_1(sK0)) [avatar definition]
195. m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (6) [avatar component clause 193] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
196. 6 [avatar split clause 113,193]
198. 7 <=> m1_subset_1(sK2,k1_zfmisc_1(sK0)) [avatar definition]
200. m1_subset_1(sK2,k1_zfmisc_1(sK0)) <- (7) [avatar component clause 198] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
201. 7 [avatar split clause 112,198]
203. 8 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
205. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (8) [avatar component clause 203] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
206. 8 [avatar split clause 111,203]
469. ~r2_hidden(X0,k1_interva1(sK0,sK1,sK2)) | r1_tarski(sK3,X0) | ~m1_subset_1(sK4,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (4) [superposition 136,185] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
470. ~r2_hidden(X0,k1_interva1(sK0,sK1,sK2)) | r1_tarski(sK3,X0) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (4, 5) [subsumption resolution 469,190] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
471. ~r2_hidden(X0,k1_interva1(sK0,sK1,sK2)) | r1_tarski(sK3,X0) <- (4, 5, 6) [subsumption resolution 470,195] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
473. ~r2_hidden(X0,k1_interva1(sK0,sK1,sK2)) | r1_tarski(X0,sK4) | ~m1_subset_1(sK4,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (4) [superposition 137,185] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
474. ~r2_hidden(X0,k1_interva1(sK0,sK1,sK2)) | r1_tarski(X0,sK4) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (4, 5) [subsumption resolution 473,190] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
475. ~r2_hidden(X0,k1_interva1(sK0,sK1,sK2)) | r1_tarski(X0,sK4) <- (4, 5, 6) [subsumption resolution 474,195] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
481. r2_hidden(sK3,k1_interva1(sK0,sK1,sK2)) | k1_xboole_0 = k1_interva1(sK0,sK1,sK2) | ~m1_subset_1(sK4,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (4) [superposition 134,185] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
484. r2_hidden(sK3,k1_interva1(sK0,sK1,sK2)) | ~m1_subset_1(sK4,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (~3, 4) [subsumption resolution 481,180] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
485. r2_hidden(sK3,k1_interva1(sK0,sK1,sK2)) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (~3, 4, 5) [subsumption resolution 484,190] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
486. r2_hidden(sK3,k1_interva1(sK0,sK1,sK2)) <- (~3, 4, 5, 6) [subsumption resolution 485,195] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
488. 40 <=> r2_hidden(sK3,k1_interva1(sK0,sK1,sK2)) [avatar definition]
490. r2_hidden(sK3,k1_interva1(sK0,sK1,sK2)) <- (40) [avatar component clause 488] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
491. 40 | 3 | ~4 | ~5 | ~6 [avatar split clause 486,193,188,183,178,488]
497. r2_hidden(sK4,k1_interva1(sK0,sK1,sK2)) | k1_xboole_0 = k1_interva1(sK0,sK1,sK2) | ~m1_subset_1(sK4,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (4) [superposition 135,185] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
500. r2_hidden(sK4,k1_interva1(sK0,sK1,sK2)) | ~m1_subset_1(sK4,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (~3, 4) [subsumption resolution 497,180] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
501. r2_hidden(sK4,k1_interva1(sK0,sK1,sK2)) | ~m1_subset_1(sK3,k1_zfmisc_1(sK0)) <- (~3, 4, 5) [subsumption resolution 500,190] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
502. r2_hidden(sK4,k1_interva1(sK0,sK1,sK2)) <- (~3, 4, 5, 6) [subsumption resolution 501,195] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
504. 41 <=> r2_hidden(sK4,k1_interva1(sK0,sK1,sK2)) [avatar definition]
506. r2_hidden(sK4,k1_interva1(sK0,sK1,sK2)) <- (41) [avatar component clause 504] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
507. 41 | 3 | ~4 | ~5 | ~6 [avatar split clause 502,193,188,183,178,504]
622. r1_tarski(sK1,sK3) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (40) [resolution 490,136] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
633. r1_tarski(sK1,sK3) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (7, 40) [subsumption resolution 622,200] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
634. r1_tarski(sK1,sK3) <- (7, 8, 40) [subsumption resolution 633,205] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
636. 55 <=> r1_tarski(sK1,sK3) [avatar definition]
638. r1_tarski(sK1,sK3) <- (55) [avatar component clause 636] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
639. 55 | ~7 | ~8 | ~40 [avatar split clause 634,488,203,198,636]
660. r1_tarski(sK4,sK2) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (41) [resolution 506,137] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
665. r1_tarski(sK4,sK2) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (7, 41) [subsumption resolution 660,200] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
666. r1_tarski(sK4,sK2) <- (7, 8, 41) [subsumption resolution 665,205] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
668. 60 <=> r1_tarski(sK4,sK2) [avatar definition]
670. r1_tarski(sK4,sK2) <- (60) [avatar component clause 668] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
671. 60 | ~7 | ~8 | ~41 [avatar split clause 666,504,203,198,668]
685. sK1 = sK3 | ~r1_tarski(sK3,sK1) <- (55) [resolution 638,141] {a:3,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
688. 63 <=> r1_tarski(sK3,sK1) [avatar definition]
716. r1_tarski(sK3,sK1) | k1_xboole_0 = k1_interva1(sK0,sK1,sK2) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (4, 5, 6) [resolution 471,134] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
722. r1_tarski(sK3,sK1) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (~3, 4, 5, 6) [subsumption resolution 716,180] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
723. r1_tarski(sK3,sK1) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (~3, 4, 5, 6, 7) [subsumption resolution 722,200] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
724. r1_tarski(sK3,sK1) <- (~3, 4, 5, 6, 7, 8) [subsumption resolution 723,205] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
725. 63 | 3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 724,203,198,193,188,183,178,688]
741. ~63 | 1 | ~55 [avatar split clause 685,636,169,688]
807. r1_tarski(sK2,sK4) | k1_xboole_0 = k1_interva1(sK0,sK1,sK2) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (4, 5, 6) [resolution 475,135] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
816. r1_tarski(sK2,sK4) | ~m1_subset_1(sK2,k1_zfmisc_1(sK0)) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (~3, 4, 5, 6) [subsumption resolution 807,180] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
817. r1_tarski(sK2,sK4) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (~3, 4, 5, 6, 7) [subsumption resolution 816,200] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
818. r1_tarski(sK2,sK4) <- (~3, 4, 5, 6, 7, 8) [subsumption resolution 817,205] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
820. 75 <=> r1_tarski(sK2,sK4) [avatar definition]
823. 75 | 3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 818,203,198,193,188,183,178,820]
832. sK2 = sK4 | ~r1_tarski(sK2,sK4) <- (60) [resolution 670,141] {a:3,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
833. ~r1_tarski(sK2,sK4) <- (~2, 60) [subsumption resolution 832,175] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
834. ~75 | 2 | ~60 [avatar split clause 833,668,173,820]
835. $false [avatar sat refutation 176,181,186,191,196,201,206,491,507,639,671,725,741,823,834] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:776,thDist:-776}
% SZS output end Proof for interva1__t6_interva1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 36
% Introduced skolems: 14
% 
% >>> Saturation
% Initial clauses: 54
% Generated clauses: 413
% Active clauses: 96
% Passive clauses: 191
% Final active clauses: 71
% Final passive clauses: 91
% Main loop iterations started: 99
% 
% >>> Simplifying Inferences
% Duplicate literals: 25
% Fw subsumption resolutions: 68
% Bw demodulations: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 1
% Forward subsumptions: 54
% Backward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 174
% Unit resulting resolution: 4
% Forward superposition: 10
% Backward superposition: 7
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 23
% Split components: 46
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 82
% SAT solver unit clauses: 11
% SAT solver binary clauses: 26
% 
% Memory used [KB]: 9850
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% forward subsumption: 0.001 s
% backward subsumption: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.005 s ( own 0.001 s ) 

