i: [1,213,0,2,1,1,0,0] t32_connsp_1
a: [3,315,0,2,1,213]
s: 315
i: [1,212,0,2,1,1,0,0] t32_connsp_1
a: [3,320,0,2,1,212]
s: 320
i: [1,211,0,2,1,1,0,0] t32_connsp_1
a: [3,325,0,2,1,211]
s: 325
i: [1,215,0,3,1,1,0,0] t32_connsp_1
a: [3,305,0,3,1,215]
s: 305
i: [1,216,0,4,1,1,0,0] t32_connsp_1
a: [3,300,0,4,1,216]
s: 300
i: [1,214,0,5,1,1,0,0] t32_connsp_1
a: [3,310,0,5,1,214]
s: 310
i: [1,217,0,2,1,0,0,255] spc1_boole
a: [3,330,0,2,1,217]
s: 330
i: [1,218,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,335,0,2,1,218]
s: 335
i: [1,219,0,2,1,0,0,2] fc1_xboole_0
a: [3,340,0,2,1,219]
s: 340
i: [1,290,0,2,1,0,0,2] rc2_xboole_0
a: [3,350,0,2,1,290]
s: 350
i: [1,291,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,355,0,2,1,291]
s: 355
i: [1,292,0,2,1,0,0,2] existence_l1_struct_0
a: [3,360,0,2,1,292]
s: 360
i: [1,293,0,2,1,0,0,2] rc1_xboole_0
a: [3,365,0,2,1,293]
s: 365
i: [1,295,0,2,1,0,0,255] rc1_relat_1
a: [3,370,0,2,1,295]
s: 370
i: [1,294,0,2,1,0,0,255] rc1_relat_1
a: [3,375,0,2,1,294]
s: 375
i: [1,296,0,2,1,0,0,255] rc2_relat_1
a: [3,380,0,2,1,296]
s: 380
i: [1,221,0,3,1,0,0,1] fc1_subset_1
s: 221
i: [1,285,0,3,1,0,0,1] rc2_subset_1
s: 285
i: [1,288,0,3,1,0,0,255] reflexivity_r1_tarski
s: 288
i: [1,220,0,3,1,0,0,255] d2_xboole_0
a: [3,345,0,3,1,220]
s: 345
r: 340
i: [1,233,0,4,2,0,0,1] dt_l1_pre_topc
s: 233
d: [2,383,1,2,1,1,40,233,315]
a: [3,388,1,2,1,383]
s: 388
i: [1,253,0,4,2,0,0,255] cc1_relat_1
s: 253
d: [2,384,1,2,1,1,40,233,355]
a: [3,393,1,2,1,384]
s: 393
i: [1,283,0,4,1,0,0,1] existence_m1_subset_1
s: 283
i: [1,287,0,4,1,0,0,255] rc3_subset_1
s: 287
d: [2,395,1,2,1,1,40,253,335]
a: [3,401,1,2,1,395]
s: 401
i: [1,223,0,5,2,0,0,255] rc6_subset_1
s: 223
d: [2,397,1,2,1,1,40,253,365]
a: [3,406,1,2,1,397]
s: 406
i: [1,227,0,5,2,0,0,1] rc1_subset_1
s: 227
i: [1,229,0,5,2,0,0,255] rc5_subset_1
s: 229
d: [2,396,1,3,1,0,40,253,285]
s: 396
i: [1,230,0,5,2,0,0,255] rc5_subset_1
s: 230
i: [1,234,0,5,2,0,0,1] fc3_struct_0
s: 234
i: [1,254,0,5,2,0,0,1] t6_boole
s: 254
i: [1,280,0,5,2,0,0,255] t7_boole
s: 280
d: [2,409,1,4,2,1,44,234,300]
d: [2,410,1,2,1,2,21,409,388]
a: [3,414,1,2,1,410]
s: 414
i: [1,284,0,5,1,0,0,1] rc2_subset_1
s: 284
d: [2,417,2,2,1,1,40,414,253]
a: [3,421,2,2,1,417]
s: 421
i: [1,286,0,5,1,0,0,255] rc3_subset_1
s: 286
d: [2,416,0,5,2,1,22,254,345]
s: 416
d: [2,408,1,5,2,0,40,234,253]
s: 408
i: [1,232,0,6,2,0,0,255] rc4_subset_1
s: 232
i: [1,235,0,6,2,0,0,1] d2_struct_0
s: 235
i: [1,237,0,6,3,0,0,255] cc7_struct_0
s: 237
i: [1,238,0,6,3,0,0,255] cc5_struct_0
s: 238
d: [2,428,1,4,2,1,40,238,360]
a: [3,434,1,2,1,428]
s: 434
i: [1,240,0,6,3,0,0,255] cc2_struct_0
s: 240
d: [2,430,2,4,2,1,40,238,393]
a: [3,443,2,2,1,430]
s: 443
i: [1,241,0,6,3,0,0,255] cc17_struct_0
s: 241
i: [1,242,0,6,3,0,0,255] cc6_struct_0
s: 242
i: [1,243,0,6,3,0,0,255] cc1_struct_0
s: 243
i: [1,246,0,6,3,0,0,255] cc4_struct_0
s: 246
i: [1,276,0,6,2,0,0,255] t1_subset
s: 276
i: [1,289,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 289
d: [2,427,0,6,2,1,22,235,345]
s: 427
d: [2,457,1,5,2,2,44,427,300]
d: [2,463,1,3,1,3,21,457,388]
a: [3,467,1,3,1,463]
s: 467
r: 421
r: 414
r: 310
r: 300
r: 305
d: [2,475,0,3,1,2,23,305,467]
a: [3,489,0,3,1,475]
s: 489
d: [2,474,0,4,1,2,23,300,467]
a: [3,484,0,4,1,474]
s: 484
d: [2,473,0,5,1,2,23,310,467]
a: [3,479,0,5,1,473]
s: 479
i: [1,222,0,7,2,0,0,255] rc6_subset_1
s: 222
d: [2,491,1,5,2,2,44,489,416]
s: 491
i: [1,226,0,7,2,0,0,1] rc1_subset_1
s: 226
d: [2,494,1,7,2,2,44,479,416]
s: 494
i: [1,228,0,7,2,0,0,255] rc5_subset_1
s: 228
i: [1,231,0,7,2,0,0,255] rc4_subset_1
s: 231
i: [1,244,0,7,3,0,0,255] cc8_struct_0
s: 244
i: [1,247,0,7,3,0,0,255] cc11_struct_0
s: 247
i: [1,248,0,7,3,0,0,255] cc11_struct_0
s: 248
i: [1,249,0,7,3,0,0,255] cc9_struct_0
s: 249
i: [1,257,0,7,3,0,0,255] fc9_struct_0
s: 257
i: [1,258,0,7,3,0,0,255] fc6_struct_0
s: 258
i: [1,260,0,7,3,0,0,255] rc22_struct_0
s: 260
i: [1,264,0,7,3,0,0,1] fc2_struct_0
s: 264
i: [1,266,0,7,3,0,0,1] rc4_struct_0
s: 266
i: [1,268,0,7,3,0,0,255] rc21_struct_0
s: 268
i: [1,269,0,7,3,0,0,255] rc21_struct_0
s: 269
i: [1,271,0,7,3,0,0,255] fc8_struct_0
s: 271
i: [1,272,0,7,3,0,0,255] fc7_struct_0
s: 272
i: [1,273,0,7,3,0,0,1] fc1_struct_0
s: 273
i: [1,278,0,7,2,0,0,255] t3_subset
s: 278
i: [1,279,0,7,3,0,0,1] t8_boole
s: 279
d: [2,495,0,7,3,1,22,244,345]
s: 495
d: [2,496,0,7,3,1,22,249,345]
s: 496
i: [1,236,0,8,2,0,0,1] dt_k1_struct_0
s: 236
i: [1,251,0,8,3,0,0,255] cc5_subset_1
s: 251
d: [2,511,2,7,3,2,40,251,494]
a: [3,525,2,3,1,511]
s: 525
i: [1,252,0,8,3,0,0,255] cc2_relat_1
s: 252
d: [2,535,3,4,2,1,40,525,272]
d: [2,536,3,2,1,2,21,535,388]
a: [3,540,3,2,1,536]
s: 540
i: [1,255,0,8,3,0,0,1] cc1_subset_1
s: 255
i: [1,277,0,8,3,0,0,255] t2_subset
s: 277
d: [2,567,1,9,2,1,40,277,479]
d: [2,578,1,5,1,1,21,567,221]
a: [3,582,1,5,1,578]
s: 582
i: [1,250,0,9,4,0,0,255] cc10_struct_0
s: 250
d: [2,590,2,5,1,1,40,582,289]
a: [3,597,2,5,1,590]
s: 597
i: [1,256,0,9,3,0,0,255] cc4_subset_1
s: 256
i: [1,263,0,9,4,0,0,255] rc7_pre_topc
s: 263
d: [2,604,2,9,3,2,40,256,494]
a: [3,618,2,3,1,604]
s: 618
i: [1,282,0,9,3,0,0,255] t5_subset
s: 282
d: [2,564,2,11,3,2,40,277,494]
d: [2,575,2,7,2,2,21,564,221]
s: 575
i: [1,259,0,10,3,0,0,255] rc22_struct_0
s: 259
i: [1,265,0,10,3,0,0,1] rc4_struct_0
s: 265
d: [2,593,2,7,2,2,44,582,416]
s: 593
i: [1,267,0,10,3,0,0,255] rc21_struct_0
s: 267
d: [2,520,1,5,2,0,40,251,284]
s: 520
i: [1,270,0,10,3,0,0,255] rc6_pre_topc
s: 270
i: [1,281,0,10,3,0,0,255] t4_subset
s: 281
d: [2,674,1,7,2,1,40,281,479]
s: 674
i: [1,224,0,11,4,0,0,255] cc3_subset_1
s: 224
r: 618
s: 618
i: [1,225,0,11,4,0,0,255] cc2_subset_1
s: 225
i: [1,262,0,12,4,0,0,255] rc7_pre_topc
s: 262
d: [2,607,1,7,2,1,40,256,479]
a: [3,625,1,4,1,607]
s: 625
i: [1,261,0,16,6,0,0,1] cc2_connsp_1
s: 261
d: [2,740,1,0,0,6,52,325,320,315,335,489,479,261]
e: 740
d: [2,744,1,11,5,1,40,261,479]
d: [2,768,1,8,4,2,21,744,489]
d: [2,769,1,6,3,3,21,768,335]
d: [2,770,1,4,2,4,21,769,315]
d: [2,771,1,2,1,5,21,770,320]
d: [2,772,1,0,0,6,21,771,325]
e: 772
d: [2,458,1,5,2,2,44,300,427]
d: [2,469,1,3,1,3,21,458,388]
f: 772,475,473,469,383,220,218,211,212,213,214,215,216
% Refutation found. Thanks to Tanya!
% SZS status Theorem for connsp_1__t32_connsp_1
% SZS output start Proof for connsp_1__t32_connsp_1
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ~(k1_struct_0(X0) = X1 & v3_connsp_1(X1,X0)))) [input t32_connsp_1]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ~(k1_struct_0(X0) = X1 & v3_connsp_1(X1,X0)))) [negated conjecture 1]
12. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
61. o_0_0_xboole_0 = k1_xboole_0 [input d2_xboole_0]
66. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
70. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v3_connsp_1(X1,X0) => ~v1_xboole_0(X1)))) [input cc2_connsp_1]
74. ! [X0] : (l1_struct_0(X0) => k1_struct_0(X0) = k1_xboole_0) [input d2_struct_0]
92. ? [X0] : (? [X1] : ((k1_struct_0(X0) = X1 & v3_connsp_1(X1,X0)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & (l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
93. ? [X0] : (? [X1] : (k1_struct_0(X0) = X1 & v3_connsp_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [flattening 92]
102. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 66]
104. ! [X0] : (k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0)) [ennf transformation 74]
140. ! [X0] : (! [X1] : ((~v1_xboole_0(X1) | ~v3_connsp_1(X1,X0)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0))) [ennf transformation 70]
141. ! [X0] : (! [X1] : (~v1_xboole_0(X1) | ~v3_connsp_1(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [flattening 140]
172. ? [X0] : (? [X1] : (k1_struct_0(X0) = X1 & v3_connsp_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (? [X1] : (k1_struct_0(sK0) = X1 & v3_connsp_1(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0)) [choice axiom]
173. ? [X1] : (k1_struct_0(sK0) = X1 & v3_connsp_1(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => (k1_struct_0(sK0) = sK1 & v3_connsp_1(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
174. (k1_struct_0(sK0) = sK1 & v3_connsp_1(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0) [skolemisation 93,173,172]
211. ~v2_struct_0(sK0) [cnf transformation 174] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
212. v2_pre_topc(sK0) [cnf transformation 174] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
213. l1_pre_topc(sK0) [cnf transformation 174] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
214. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 174] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
215. v3_connsp_1(sK1,sK0) [cnf transformation 174] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
216. k1_struct_0(sK0) = sK1 [cnf transformation 174] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
218. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 12] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
220. o_0_0_xboole_0 = k1_xboole_0 [cnf transformation 61] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
233. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 102] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
235. k1_struct_0(X0) = k1_xboole_0 | ~l1_struct_0(X0) [cnf transformation 104] {a:0,w:6,wCS:2400,sine:1,thAx:0,allAx:1,thDist:-1}
261. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v3_connsp_1(X1,X0) | ~v1_xboole_0(X1) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 141] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
298. 1 <=> k1_struct_0(sK0) = sK1 [avatar definition]
300. k1_struct_0(sK0) = sK1 <- (1) [avatar component clause 298] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
301. 1 [avatar split clause 216,298]
303. 2 <=> v3_connsp_1(sK1,sK0) [avatar definition]
305. v3_connsp_1(sK1,sK0) <- (2) [avatar component clause 303] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
306. 2 [avatar split clause 215,303]
308. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
310. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (3) [avatar component clause 308] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
311. 3 [avatar split clause 214,308]
313. 4 <=> l1_pre_topc(sK0) [avatar definition]
315. l1_pre_topc(sK0) <- (4) [avatar component clause 313] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
316. 4 [avatar split clause 213,313]
318. 5 <=> v2_pre_topc(sK0) [avatar definition]
320. v2_pre_topc(sK0) <- (5) [avatar component clause 318] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
321. 5 [avatar split clause 212,318]
323. 6 <=> v2_struct_0(sK0) [avatar definition]
325. ~v2_struct_0(sK0) <- (~6) [avatar component clause 323] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
326. ~6 [avatar split clause 211,323]
333. 8 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
335. v1_xboole_0(o_0_0_xboole_0) <- (8) [avatar component clause 333] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
336. 8 [avatar split clause 218,333]
343. 10 <=> o_0_0_xboole_0 = k1_xboole_0 [avatar definition]
345. o_0_0_xboole_0 = k1_xboole_0 <- (10) [avatar component clause 343] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
346. 10 [avatar split clause 220,343]
383. l1_struct_0(sK0) <- (4) [resolution 233,315] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
386. 18 <=> l1_struct_0(sK0) [avatar definition]
388. l1_struct_0(sK0) <- (18) [avatar component clause 386] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
389. 18 | ~4 [avatar split clause 383,313,386]
427. k1_struct_0(X0) = o_0_0_xboole_0 | ~l1_struct_0(X0) <- (10) [forward demodulation 235,345] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
458. o_0_0_xboole_0 = sK1 | ~l1_struct_0(sK0) <- (1, 10) [superposition 300,427] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
465. 29 <=> o_0_0_xboole_0 = sK1 [avatar definition]
467. o_0_0_xboole_0 = sK1 <- (29) [avatar component clause 465] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
469. o_0_0_xboole_0 = sK1 <- (1, 10, 18) [subsumption resolution 458,388] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
470. 29 | ~1 | ~10 | ~18 [avatar split clause 469,386,343,298,465]
473. m1_subset_1(o_0_0_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) <- (3, 29) [backward demodulation 310,467] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
475. v3_connsp_1(o_0_0_xboole_0,sK0) <- (2, 29) [backward demodulation 305,467] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
477. 30 <=> m1_subset_1(o_0_0_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
479. m1_subset_1(o_0_0_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) <- (30) [avatar component clause 477] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
480. 30 | ~3 | ~29 [avatar split clause 473,465,308,477]
487. 32 <=> v3_connsp_1(o_0_0_xboole_0,sK0) [avatar definition]
489. v3_connsp_1(o_0_0_xboole_0,sK0) <- (32) [avatar component clause 487] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
490. 32 | ~2 | ~29 [avatar split clause 475,465,303,487]
744. ~v3_connsp_1(o_0_0_xboole_0,sK0) | ~v1_xboole_0(o_0_0_xboole_0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (30) [resolution 261,479] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
768. ~v1_xboole_0(o_0_0_xboole_0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (30, 32) [subsumption resolution 744,489] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
769. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (8, 30, 32) [subsumption resolution 768,335] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
770. ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 8, 30, 32) [subsumption resolution 769,315] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
771. v2_struct_0(sK0) <- (4, 5, 8, 30, 32) [subsumption resolution 770,320] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
772. $false <- (4, 5, ~6, 8, 30, 32) [subsumption resolution 771,325] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
773. ~4 | ~5 | 6 | ~8 | ~30 | ~32 [avatar contradiction clause 772]
777. $false [avatar sat refutation 301,306,311,316,321,326,336,346,389,470,480,490,773] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:264,thDist:-264}
% SZS output end Proof for connsp_1__t32_connsp_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 73
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 5
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 86
% Generated clauses: 395
% Active clauses: 103
% Passive clauses: 206
% Final active clauses: 96
% Final passive clauses: 89
% Main loop iterations started: 113
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 40
% Fw demodulations: 4
% Bw demodulations: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 22
% Forward subsumptions: 59
% 
% >>> Generating Inferences
% Binary resolution: 174
% Unit resulting resolution: 1
% Forward superposition: 12
% Backward superposition: 7
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 11
% Split components: 22
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 48
% SAT solver unit clauses: 17
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9722
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% sat proof minimization: 0.001 s
% resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

