i: [1,127,0,3,1,1,0,0] l49_mod_2
a: [3,231,0,3,1,127]
s: 231
i: [1,128,0,22,3,1,0,0] l49_mod_2
a: [3,222,0,5,1,128]
s: 222
d: [2,389,0,5,1,2,22,222,231]
a: [3,393,0,5,1,389]
s: 393
i: [1,126,0,6,1,1,0,0] l49_mod_2
a: [3,236,0,6,1,126]
s: 236
i: [1,125,0,6,1,1,0,0] l49_mod_2
a: [3,241,0,6,1,125]
s: 241
i: [1,124,0,6,1,1,0,0] l49_mod_2
a: [3,246,0,6,1,124]
s: 246
i: [1,123,0,6,1,1,0,0] l49_mod_2
a: [3,251,0,6,1,123]
s: 251
i: [1,129,0,2,1,0,0,1] spc2_boole
a: [3,256,0,2,1,129]
s: 256
i: [1,130,0,2,1,0,0,1] spc1_boole
a: [3,261,0,2,1,130]
s: 261
i: [1,131,0,2,1,0,0,255] fc1_xboole_0
a: [3,266,0,2,1,131]
s: 266
i: [1,191,0,2,1,0,0,2] rc2_xboole_0
a: [3,381,0,2,1,191]
s: 381
i: [1,192,0,2,1,0,0,2] rc1_xboole_0
a: [3,386,0,2,1,192]
s: 386
i: [1,143,0,3,1,0,0,3] fc1_subset_1
s: 143
i: [1,180,0,3,1,0,0,3] rc2_subset_1
s: 180
i: [1,183,0,3,1,0,0,255] reflexivity_r1_tarski
s: 183
i: [1,132,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,271,0,3,1,132]
s: 271
i: [1,133,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,276,0,3,1,133]
s: 276
r: 266
d: [2,397,0,2,1,2,23,266,276]
a: [3,401,0,2,1,397]
s: 401
i: [1,134,0,3,1,0,0,255] t1_numerals
a: [3,281,0,3,1,134]
s: 281
i: [1,139,0,3,1,0,0,255] spc2_numerals
a: [3,296,0,3,1,139]
s: 296
i: [1,138,0,3,1,0,0,255] spc2_numerals
a: [3,301,0,3,1,138]
s: 301
i: [1,142,0,3,1,0,0,255] spc1_numerals
a: [3,311,0,3,1,142]
s: 311
i: [1,141,0,3,1,0,0,255] spc1_numerals
a: [3,316,0,3,1,141]
s: 316
d: [2,403,0,3,1,2,22,281,271]
d: [2,404,0,3,1,3,22,403,276]
a: [3,408,0,3,1,404]
s: 408
i: [1,178,0,4,1,0,0,1] existence_m1_subset_1
s: 178
i: [1,182,0,4,1,0,0,255] rc3_subset_1
s: 182
i: [1,135,0,4,1,0,0,255] dt_k5_numbers
a: [3,286,0,4,1,135]
s: 286
i: [1,136,0,4,1,0,0,255] dt_k6_numbers
a: [3,291,0,4,1,136]
s: 291
i: [1,137,0,4,1,0,0,255] spc2_numerals
a: [3,306,0,4,1,137]
s: 306
i: [1,140,0,4,1,0,0,255] spc1_numerals
a: [3,321,0,4,1,140]
s: 321
i: [1,195,0,10,2,0,0,1] d12_mod_2
a: [3,348,0,4,1,195]
s: 348
i: [1,147,0,5,2,0,0,3] rc1_subset_1
s: 147
i: [1,150,0,5,2,0,0,255] t6_boole
s: 150
i: [1,167,0,5,2,0,0,2] t7_boole
s: 167
i: [1,179,0,5,1,0,0,3] rc2_subset_1
s: 179
i: [1,181,0,5,1,0,0,255] rc3_subset_1
s: 181
i: [1,190,0,5,1,0,0,1] fc2_subset_1
s: 190
d: [2,410,0,5,2,1,22,150,276]
s: 410
d: [2,417,1,7,2,2,44,393,410]
s: 417
i: [1,149,0,6,2,0,0,255] rc4_subset_1
s: 149
i: [1,163,0,6,2,0,0,2] t1_subset
s: 163
d: [2,412,1,8,2,2,44,241,410]
s: 412
i: [1,184,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 184
d: [2,413,1,8,2,2,44,246,410]
s: 413
i: [1,207,0,6,1,0,0,1] d1_enumset1
s: 207
i: [1,209,0,6,1,0,0,1] d1_enumset1
s: 209
d: [2,414,1,8,2,2,44,251,410]
s: 414
i: [1,211,0,6,1,0,0,1] d1_enumset1
s: 211
d: [2,419,1,5,2,2,44,408,410]
s: 419
i: [1,193,0,10,2,0,0,1] d12_mod_2
a: [3,326,0,6,1,193]
s: 326
d: [2,429,1,0,0,1,52,207,326,163]
e: 429
r: 326
a: [3,330,0,4,1,193]
s: 330
i: [1,197,0,17,3,0,0,1] d13_mod_2
d: [2,214,0,11,2,0,18,197]
a: [3,353,0,5,1,214]
s: 353
a: [3,325,0,6,1,193]
s: 325
i: [1,194,0,10,2,0,0,1] d12_mod_2
a: [3,335,0,6,1,194]
s: 335
d: [2,433,1,0,0,1,52,209,335,163]
e: 433
r: 335
a: [3,339,0,4,1,194]
s: 339
i: [1,199,0,17,3,0,0,1] d13_mod_2
a: [3,358,0,5,1,199]
s: 358
i: [1,201,0,17,3,0,0,1] d13_mod_2
a: [3,363,0,5,1,201]
s: 363
i: [1,203,0,17,3,0,0,1] d13_mod_2
d: [2,213,0,11,2,0,18,203]
a: [3,368,0,5,1,213]
s: 368
a: [3,334,0,6,1,194]
s: 334
a: [3,343,0,6,1,195]
s: 343
i: [1,146,0,7,2,0,0,3] rc1_subset_1
s: 146
d: [2,415,1,6,2,2,44,291,410]
s: 415
i: [1,148,0,7,2,0,0,255] rc4_subset_1
s: 148
d: [2,416,1,6,2,2,44,348,410]
s: 416
i: [1,165,0,7,2,0,0,255] t3_subset
s: 165
i: [1,166,0,7,3,0,0,1] t8_boole
s: 166
d: [2,420,1,6,1,0,40,207,184]
s: 420
i: [1,151,0,8,3,0,0,3] cc1_subset_1
s: 151
d: [2,443,1,4,2,1,40,151,286]
a: [3,452,1,2,1,443]
s: 452
i: [1,164,0,8,3,0,0,2] t2_subset
s: 164
r: 452
i: [1,152,0,9,3,0,0,255] cc4_subset_1
s: 152
d: [2,462,1,11,2,1,40,164,251]
d: [2,480,1,6,1,1,21,462,190]
a: [3,484,1,6,1,480]
s: 484
i: [1,169,0,9,3,0,0,2] t5_subset
s: 169
d: [2,463,1,11,2,1,40,164,246]
d: [2,486,1,6,1,1,21,463,190]
a: [3,490,1,6,1,486]
s: 490
i: [1,168,0,10,3,0,0,2] t4_subset
s: 168
i: [1,144,0,11,4,0,0,255] cc3_subset_1
s: 144
d: [2,464,1,11,2,1,40,164,241]
d: [2,492,1,6,1,1,21,464,190]
a: [3,496,1,6,1,492]
s: 496
i: [1,145,0,11,4,0,0,255] cc2_subset_1
s: 145
d: [2,557,2,6,1,1,40,484,184]
a: [3,564,2,6,1,557]
s: 564
i: [1,204,0,17,3,0,0,1] d13_mod_2
a: [3,372,0,11,2,204]
s: 372
d: [2,615,1,0,0,3,52,393,251,372]
e: 615
d: [2,619,1,5,1,2,40,372,251]
d: [2,628,1,0,0,3,21,619,393]
e: 628
d: [2,624,2,7,2,3,40,372,414]
d: [2,632,2,2,1,4,21,624,393]
d: [2,633,2,0,0,5,21,632,401]
e: 633
r: 417
r: 393
a: [3,392,0,5,1,389]
s: 392
d: [2,623,2,7,2,3,40,372,413]
d: [2,631,2,5,1,4,21,623,401]
a: [3,654,2,5,1,631]
s: 654
a: [3,218,0,6,1,128]
d: [2,640,0,6,1,2,22,218,231]
a: [3,644,0,6,1,640]
s: 644
d: [2,680,2,7,2,3,44,644,416]
a: [3,684,2,2,1,680]
s: 684
i: [1,205,0,17,3,0,0,1] d13_mod_2
a: [3,376,0,11,2,205]
s: 376
i: [1,153,0,13,2,0,0,1] dt_k10_mod_2
s: 153
d: [2,693,1,5,1,2,40,376,251]
a: [3,715,1,5,1,693]
s: 715
i: [1,188,0,14,4,0,0,255] existence_m2_subset_1
s: 188
d: [2,694,1,5,1,2,40,376,246]
a: [3,720,1,5,1,694]
s: 720
d: [2,748,0,12,3,0,21,188,151]
s: 748
i: [1,185,0,15,5,0,0,255] dt_m2_subset_1
s: 185
d: [2,695,1,5,1,2,40,376,241]
a: [3,725,1,5,1,695]
s: 725
d: [2,750,0,13,4,0,21,185,151]
s: 750
d: [2,622,2,7,2,3,40,372,412]
d: [2,630,2,5,1,4,21,622,401]
a: [3,659,2,5,1,630]
s: 659
i: [1,186,0,15,5,0,0,255] redefinition_m2_subset_1
s: 186
d: [2,770,0,13,4,0,21,186,151]
s: 770
d: [2,577,2,6,1,1,40,490,184]
a: [3,584,2,6,1,577]
s: 584
i: [1,187,0,15,5,0,0,255] redefinition_m2_subset_1
s: 187
d: [2,605,2,6,1,1,40,496,184]
a: [3,612,2,6,1,605]
s: 612
d: [2,778,0,13,4,0,21,187,151]
s: 778
i: [1,212,0,15,4,0,0,1] d1_enumset1
s: 212
d: [2,395,0,6,1,2,22,236,231]
d: [2,620,1,5,1,2,40,372,246]
d: [2,618,1,5,1,2,40,372,325]
d: [2,617,1,5,1,2,40,372,334]
d: [2,691,1,5,1,2,40,376,334]
d: [2,692,1,5,1,2,40,376,325]
d: [2,795,2,9,3,1,40,212,484]
d: [2,796,2,9,3,1,40,212,490]
d: [2,797,2,9,3,1,40,212,496]
f: 797,796,795,695,694,692,691,680,617,618,620,640,619,492,486,480,433,429,395,389,205,204,213,201,199,214,195,194,193,133,131,123,124,125,126,127,128
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mod_2__l49_mod_2
% SZS output start Proof for mod_2__l49_mod_2
1. ! [X0] : (m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) => ! [X1] : (m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) => ! [X2] : (m1_subset_1(X2,k1_enumset1(k6_numbers,1,2)) => ! [X3] : (m1_subset_1(X3,k1_enumset1(k6_numbers,1,2)) => (k6_numbers = X3 => (k11_mod_2(k11_mod_2(X0,X1),X2) = k11_mod_2(X0,k11_mod_2(X1,X2)) & k11_mod_2(X0,X3) = X0 & k11_mod_2(X0,k10_mod_2(X0)) = X3)))))) [input l49_mod_2]
2. ~! [X0] : (m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) => ! [X1] : (m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) => ! [X2] : (m1_subset_1(X2,k1_enumset1(k6_numbers,1,2)) => ! [X3] : (m1_subset_1(X3,k1_enumset1(k6_numbers,1,2)) => (k6_numbers = X3 => (k11_mod_2(k11_mod_2(X0,X1),X2) = k11_mod_2(X0,k11_mod_2(X1,X2)) & k11_mod_2(X0,X3) = X0 & k11_mod_2(X0,k10_mod_2(X0)) = X3)))))) [negated conjecture 1]
8. ! [X0] : (m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) => ((2 = X0 => 1 = k10_mod_2(X0)) & (1 = X0 => 2 = k10_mod_2(X0)) & (k6_numbers = X0 => k6_numbers = k10_mod_2(X0)))) [input d12_mod_2]
9. ! [X0] : (m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) => ! [X1] : (m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) => (((2 = X1 & 2 = X0) => 1 = k11_mod_2(X0,X1)) & ((1 = X1 & 2 = X0) => k6_numbers = k11_mod_2(X0,X1)) & ((2 = X1 & 1 = X0) => k6_numbers = k11_mod_2(X0,X1)) & ((1 = X1 & 1 = X0) => 2 = k11_mod_2(X0,X1)) & (k6_numbers = X1 => k11_mod_2(X0,X1) = X0) & (k6_numbers = X0 => k11_mod_2(X0,X1) = X1)))) [input d13_mod_2]
10. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (k1_enumset1(X0,X1,X2) = X3 <=> ! [X4] : (r2_hidden(X4,X3) <=> ~(X2 != X4 & X1 != X4 & X0 != X4))) [input d1_enumset1]
25. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
26. ! [X0,X1,X2] : ~v1_xboole_0(k1_enumset1(X0,X1,X2)) [input fc2_subset_1]
34. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
42. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
43. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
47. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
50. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 42]
51. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 43]
57. ! [X0,X1,X2,X3] : (k1_enumset1(X0,X1,X2) = X3 <=> ! [X4] : (r2_hidden(X4,X3) <=> ~(X2 != X4 & X1 != X4 & X0 != X4))) [flattening 10]
62. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((k11_mod_2(k11_mod_2(X0,X1),X2) != k11_mod_2(X0,k11_mod_2(X1,X2)) | k11_mod_2(X0,X3) != X0 | k11_mod_2(X0,k10_mod_2(X0)) != X3) & k6_numbers = X3) & m1_subset_1(X3,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X2,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X1,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X0,k1_enumset1(k6_numbers,1,2))) [ennf transformation 2]
63. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((k11_mod_2(k11_mod_2(X0,X1),X2) != k11_mod_2(X0,k11_mod_2(X1,X2)) | k11_mod_2(X0,X3) != X0 | k11_mod_2(X0,k10_mod_2(X0)) != X3) & k6_numbers = X3 & m1_subset_1(X3,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X2,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X1,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X0,k1_enumset1(k6_numbers,1,2))) [flattening 62]
70. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 47]
74. ! [X0] : (((1 = k10_mod_2(X0) | 2 != X0) & (2 = k10_mod_2(X0) | 1 != X0) & (k6_numbers = k10_mod_2(X0) | k6_numbers != X0)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2))) [ennf transformation 8]
75. ! [X0] : (! [X1] : (((1 = k11_mod_2(X0,X1) | (2 != X1 | 2 != X0)) & (k6_numbers = k11_mod_2(X0,X1) | (1 != X1 | 2 != X0)) & (k6_numbers = k11_mod_2(X0,X1) | (2 != X1 | 1 != X0)) & (2 = k11_mod_2(X0,X1) | (1 != X1 | 1 != X0)) & (k11_mod_2(X0,X1) = X0 | k6_numbers != X1) & (k11_mod_2(X0,X1) = X1 | k6_numbers != X0)) | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2))) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2))) [ennf transformation 9]
76. ! [X0] : (! [X1] : (((1 = k11_mod_2(X0,X1) | 2 != X1 | 2 != X0) & (k6_numbers = k11_mod_2(X0,X1) | 1 != X1 | 2 != X0) & (k6_numbers = k11_mod_2(X0,X1) | 2 != X1 | 1 != X0) & (2 = k11_mod_2(X0,X1) | 1 != X1 | 1 != X0) & (k11_mod_2(X0,X1) = X0 | k6_numbers != X1) & (k11_mod_2(X0,X1) = X1 | k6_numbers != X0)) | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2))) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2))) [flattening 75]
77. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 50]
78. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 51]
79. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 78]
86. ! [X0,X1,X2,X3] : (k1_enumset1(X0,X1,X2) = X3 <=> ! [X4] : (r2_hidden(X4,X3) <=> (X2 = X4 | X1 = X4 | X0 = X4))) [ennf transformation 57]
96. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((k11_mod_2(k11_mod_2(X0,X1),X2) != k11_mod_2(X0,k11_mod_2(X1,X2)) | k11_mod_2(X0,X3) != X0 | k11_mod_2(X0,k10_mod_2(X0)) != X3) & k6_numbers = X3 & m1_subset_1(X3,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X2,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X1,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X0,k1_enumset1(k6_numbers,1,2))) => (? [X1] : (? [X2] : (? [X3] : ((k11_mod_2(k11_mod_2(sK0,X1),X2) != k11_mod_2(sK0,k11_mod_2(X1,X2)) | sK0 != k11_mod_2(sK0,X3) | k11_mod_2(sK0,k10_mod_2(sK0)) != X3) & k6_numbers = X3 & m1_subset_1(X3,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X2,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X1,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(sK0,k1_enumset1(k6_numbers,1,2))) [choice axiom]
97. ? [X1] : (? [X2] : (? [X3] : ((k11_mod_2(k11_mod_2(sK0,X1),X2) != k11_mod_2(sK0,k11_mod_2(X1,X2)) | sK0 != k11_mod_2(sK0,X3) | k11_mod_2(sK0,k10_mod_2(sK0)) != X3) & k6_numbers = X3 & m1_subset_1(X3,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X2,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X1,k1_enumset1(k6_numbers,1,2))) => (? [X2] : (? [X3] : ((k11_mod_2(k11_mod_2(sK0,sK1),X2) != k11_mod_2(sK0,k11_mod_2(sK1,X2)) | sK0 != k11_mod_2(sK0,X3) | k11_mod_2(sK0,k10_mod_2(sK0)) != X3) & k6_numbers = X3 & m1_subset_1(X3,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X2,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(sK1,k1_enumset1(k6_numbers,1,2))) [choice axiom]
98. ? [X2] : (? [X3] : ((k11_mod_2(k11_mod_2(sK0,sK1),X2) != k11_mod_2(sK0,k11_mod_2(sK1,X2)) | sK0 != k11_mod_2(sK0,X3) | k11_mod_2(sK0,k10_mod_2(sK0)) != X3) & k6_numbers = X3 & m1_subset_1(X3,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(X2,k1_enumset1(k6_numbers,1,2))) => (? [X3] : ((k11_mod_2(k11_mod_2(sK0,sK1),sK2) != k11_mod_2(sK0,k11_mod_2(sK1,sK2)) | sK0 != k11_mod_2(sK0,X3) | k11_mod_2(sK0,k10_mod_2(sK0)) != X3) & k6_numbers = X3 & m1_subset_1(X3,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(sK2,k1_enumset1(k6_numbers,1,2))) [choice axiom]
99. ? [X3] : ((k11_mod_2(k11_mod_2(sK0,sK1),sK2) != k11_mod_2(sK0,k11_mod_2(sK1,sK2)) | sK0 != k11_mod_2(sK0,X3) | k11_mod_2(sK0,k10_mod_2(sK0)) != X3) & k6_numbers = X3 & m1_subset_1(X3,k1_enumset1(k6_numbers,1,2))) => ((k11_mod_2(k11_mod_2(sK0,sK1),sK2) != k11_mod_2(sK0,k11_mod_2(sK1,sK2)) | sK0 != k11_mod_2(sK0,sK3) | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3) & k6_numbers = sK3 & m1_subset_1(sK3,k1_enumset1(k6_numbers,1,2))) [choice axiom]
100. ((((k11_mod_2(k11_mod_2(sK0,sK1),sK2) != k11_mod_2(sK0,k11_mod_2(sK1,sK2)) | sK0 != k11_mod_2(sK0,sK3) | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3) & k6_numbers = sK3 & m1_subset_1(sK3,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(sK2,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(sK1,k1_enumset1(k6_numbers,1,2))) & m1_subset_1(sK0,k1_enumset1(k6_numbers,1,2)) [skolemisation 63,99,98,97,96]
105. ! [X0,X1,X2,X3] : ((k1_enumset1(X0,X1,X2) = X3 | ? [X4] : (((X2 != X4 & X1 != X4 & X0 != X4) | ~r2_hidden(X4,X3)) & ((X2 = X4 | X1 = X4 | X0 = X4) | r2_hidden(X4,X3)))) & (! [X4] : ((r2_hidden(X4,X3) | (X2 != X4 & X1 != X4 & X0 != X4)) & ((X2 = X4 | X1 = X4 | X0 = X4) | ~r2_hidden(X4,X3))) | k1_enumset1(X0,X1,X2) != X3)) [nnf transformation 86]
106. ! [X0,X1,X2,X3] : ((k1_enumset1(X0,X1,X2) = X3 | ? [X4] : (((X2 != X4 & X1 != X4 & X0 != X4) | ~r2_hidden(X4,X3)) & (X2 = X4 | X1 = X4 | X0 = X4 | r2_hidden(X4,X3)))) & (! [X4] : ((r2_hidden(X4,X3) | (X2 != X4 & X1 != X4 & X0 != X4)) & (X2 = X4 | X1 = X4 | X0 = X4 | ~r2_hidden(X4,X3))) | k1_enumset1(X0,X1,X2) != X3)) [flattening 105]
107. ! [X0,X1,X2,X3] : ((k1_enumset1(X0,X1,X2) = X3 | ? [X4] : (((X2 != X4 & X1 != X4 & X0 != X4) | ~r2_hidden(X4,X3)) & (X2 = X4 | X1 = X4 | X0 = X4 | r2_hidden(X4,X3)))) & (! [X5] : ((r2_hidden(X5,X3) | (X2 != X5 & X1 != X5 & X0 != X5)) & (X2 = X5 | X1 = X5 | X0 = X5 | ~r2_hidden(X5,X3))) | k1_enumset1(X0,X1,X2) != X3)) [rectify 106]
108. ! [X3,X2,X1,X0] : (? [X4] : (((X2 != X4 & X1 != X4 & X0 != X4) | ~r2_hidden(X4,X3)) & (X2 = X4 | X1 = X4 | X0 = X4 | r2_hidden(X4,X3))) => (((sK6(X0,X1,X2,X3) != X2 & sK6(X0,X1,X2,X3) != X1 & sK6(X0,X1,X2,X3) != X0) | ~r2_hidden(sK6(X0,X1,X2,X3),X3)) & (sK6(X0,X1,X2,X3) = X2 | sK6(X0,X1,X2,X3) = X1 | sK6(X0,X1,X2,X3) = X0 | r2_hidden(sK6(X0,X1,X2,X3),X3)))) [choice axiom]
109. ! [X0,X1,X2,X3] : ((k1_enumset1(X0,X1,X2) = X3 | (((sK6(X0,X1,X2,X3) != X2 & sK6(X0,X1,X2,X3) != X1 & sK6(X0,X1,X2,X3) != X0) | ~r2_hidden(sK6(X0,X1,X2,X3),X3)) & (sK6(X0,X1,X2,X3) = X2 | sK6(X0,X1,X2,X3) = X1 | sK6(X0,X1,X2,X3) = X0 | r2_hidden(sK6(X0,X1,X2,X3),X3)))) & (! [X5] : ((r2_hidden(X5,X3) | (X2 != X5 & X1 != X5 & X0 != X5)) & (X2 = X5 | X1 = X5 | X0 = X5 | ~r2_hidden(X5,X3))) | k1_enumset1(X0,X1,X2) != X3)) [skolemisation 107,108]
123. m1_subset_1(sK0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 100] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
124. m1_subset_1(sK1,k1_enumset1(k6_numbers,1,2)) [cnf transformation 100] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
125. m1_subset_1(sK2,k1_enumset1(k6_numbers,1,2)) [cnf transformation 100] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
126. m1_subset_1(sK3,k1_enumset1(k6_numbers,1,2)) [cnf transformation 100] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
127. k6_numbers = sK3 [cnf transformation 100] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
128. k11_mod_2(k11_mod_2(sK0,sK1),sK2) != k11_mod_2(sK0,k11_mod_2(sK1,sK2)) | sK0 != k11_mod_2(sK0,sK3) | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 [cnf transformation 100] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
131. v1_xboole_0(k1_xboole_0) [cnf transformation 25] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
133. k6_numbers = k1_xboole_0 [cnf transformation 34] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
150. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 70] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
154. k6_numbers = k10_mod_2(X0) | k6_numbers != X0 | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 74] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:1,thDist:-1}
155. 2 = k10_mod_2(X0) | 1 != X0 | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 74] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:1,thDist:-1}
156. 1 = k10_mod_2(X0) | 2 != X0 | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 74] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:1,thDist:-1}
157. k11_mod_2(X0,X1) = X1 | k6_numbers != X0 | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 76] {a:0,w:20,wCS:8000,sine:1,thAx:0,allAx:1,thDist:-1}
158. k11_mod_2(X0,X1) = X0 | k6_numbers != X1 | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 76] {a:0,w:20,wCS:8000,sine:1,thAx:0,allAx:1,thDist:-1}
159. 2 = k11_mod_2(X0,X1) | 1 != X1 | 1 != X0 | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 76] {a:0,w:23,wCS:9200,sine:1,thAx:0,allAx:1,thDist:-1}
160. k6_numbers = k11_mod_2(X0,X1) | 2 != X1 | 1 != X0 | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 76] {a:0,w:23,wCS:9200,sine:1,thAx:0,allAx:1,thDist:-1}
161. k6_numbers = k11_mod_2(X0,X1) | 1 != X1 | 2 != X0 | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 76] {a:0,w:23,wCS:9200,sine:1,thAx:0,allAx:1,thDist:-1}
162. 1 = k11_mod_2(X0,X1) | 2 != X1 | 2 != X0 | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [cnf transformation 76] {a:0,w:23,wCS:9200,sine:1,thAx:0,allAx:1,thDist:-1}
163. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 77] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
164. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 79] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
170. X2 = X5 | X1 = X5 | X0 = X5 | ~r2_hidden(X5,X3) | k1_enumset1(X0,X1,X2) != X3 [cnf transformation 109] {a:0,w:18,wCS:7200,sine:1,thAx:0,allAx:2,thDist:-2}
172. r2_hidden(X5,X3) | X1 != X5 | k1_enumset1(X0,X1,X2) != X3 [cnf transformation 109] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
173. r2_hidden(X5,X3) | X2 != X5 | k1_enumset1(X0,X1,X2) != X3 [cnf transformation 109] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
190. ~v1_xboole_0(k1_enumset1(X0,X1,X2)) [cnf transformation 26] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
193. 1 = k10_mod_2(2) | ~m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) [equality resolution 156] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:1,thDist:-1}
194. 2 = k10_mod_2(1) | ~m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) [equality resolution 155] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:1,thDist:-1}
195. k6_numbers = k10_mod_2(k6_numbers) | ~m1_subset_1(k6_numbers,k1_enumset1(k6_numbers,1,2)) [equality resolution 154] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:1,thDist:-1}
196. 1 = k11_mod_2(X0,2) | 2 != X0 | ~m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [equality resolution 162] {a:0,w:20,wCS:8000,sine:1,thAx:0,allAx:1,thDist:-1}
197. 1 = k11_mod_2(2,2) | ~m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) [equality resolution 196] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
198. k6_numbers = k11_mod_2(X0,1) | 2 != X0 | ~m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [equality resolution 161] {a:0,w:20,wCS:8000,sine:1,thAx:0,allAx:1,thDist:-1}
199. k6_numbers = k11_mod_2(2,1) | ~m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) [equality resolution 198] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
200. k6_numbers = k11_mod_2(X0,2) | 1 != X0 | ~m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [equality resolution 160] {a:0,w:20,wCS:8000,sine:1,thAx:0,allAx:1,thDist:-1}
201. k6_numbers = k11_mod_2(1,2) | ~m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) [equality resolution 200] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
202. 2 = k11_mod_2(X0,1) | 1 != X0 | ~m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [equality resolution 159] {a:0,w:20,wCS:8000,sine:1,thAx:0,allAx:1,thDist:-1}
203. 2 = k11_mod_2(1,1) | ~m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) [equality resolution 202] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
204. k11_mod_2(X0,k6_numbers) = X0 | ~m1_subset_1(k6_numbers,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) [equality resolution 158] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
205. k11_mod_2(k6_numbers,X1) = X1 | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) | ~m1_subset_1(k6_numbers,k1_enumset1(k6_numbers,1,2)) [equality resolution 157] {a:0,w:17,wCS:6800,sine:1,thAx:0,allAx:1,thDist:-1}
206. r2_hidden(X5,X3) | k1_enumset1(X0,X1,X5) != X3 [equality resolution 173] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
207. r2_hidden(X5,k1_enumset1(X0,X1,X5)) [equality resolution 206] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
208. r2_hidden(X5,X3) | k1_enumset1(X0,X5,X2) != X3 [equality resolution 172] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
209. r2_hidden(X5,k1_enumset1(X0,X5,X2)) [equality resolution 208] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
212. ~r2_hidden(X5,k1_enumset1(X0,X1,X2)) | X1 = X5 | X0 = X5 | X2 = X5 [equality resolution 170] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
213. 2 = k11_mod_2(1,1) | ~m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) [duplicate literal removal 203] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:1,thDist:-1}
214. 1 = k11_mod_2(2,2) | ~m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) [duplicate literal removal 197] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:1,thDist:-1}
216. 1 <=> k11_mod_2(sK0,k10_mod_2(sK0)) = sK3 [avatar definition]
218. k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 <- (~1) [avatar component clause 216] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
220. 2 <=> sK0 = k11_mod_2(sK0,sK3) [avatar definition]
222. sK0 != k11_mod_2(sK0,sK3) <- (~2) [avatar component clause 220] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
224. 3 <=> k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [avatar definition]
227. ~1 | ~2 | ~3 [avatar split clause 128,224,220,216]
229. 4 <=> k6_numbers = sK3 [avatar definition]
231. k6_numbers = sK3 <- (4) [avatar component clause 229] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
232. 4 [avatar split clause 127,229]
234. 5 <=> m1_subset_1(sK3,k1_enumset1(k6_numbers,1,2)) [avatar definition]
236. m1_subset_1(sK3,k1_enumset1(k6_numbers,1,2)) <- (5) [avatar component clause 234] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
237. 5 [avatar split clause 126,234]
239. 6 <=> m1_subset_1(sK2,k1_enumset1(k6_numbers,1,2)) [avatar definition]
241. m1_subset_1(sK2,k1_enumset1(k6_numbers,1,2)) <- (6) [avatar component clause 239] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
242. 6 [avatar split clause 125,239]
244. 7 <=> m1_subset_1(sK1,k1_enumset1(k6_numbers,1,2)) [avatar definition]
246. m1_subset_1(sK1,k1_enumset1(k6_numbers,1,2)) <- (7) [avatar component clause 244] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
247. 7 [avatar split clause 124,244]
249. 8 <=> m1_subset_1(sK0,k1_enumset1(k6_numbers,1,2)) [avatar definition]
251. m1_subset_1(sK0,k1_enumset1(k6_numbers,1,2)) <- (8) [avatar component clause 249] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
252. 8 [avatar split clause 123,249]
264. 11 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
267. 11 [avatar split clause 131,264]
274. 13 <=> k6_numbers = k1_xboole_0 [avatar definition]
276. k6_numbers = k1_xboole_0 <- (13) [avatar component clause 274] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
277. 13 [avatar split clause 133,274]
324. 23 <=> m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) [avatar definition]
325. m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) <- (23) [avatar component clause 324] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
326. ~m1_subset_1(2,k1_enumset1(k6_numbers,1,2)) <- (~23) [avatar component clause 324] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
328. 24 <=> 1 = k10_mod_2(2) [avatar definition]
331. ~23 | 24 [avatar split clause 193,328,324]
333. 25 <=> m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) [avatar definition]
334. m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) <- (25) [avatar component clause 333] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
335. ~m1_subset_1(1,k1_enumset1(k6_numbers,1,2)) <- (~25) [avatar component clause 333] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
337. 26 <=> 2 = k10_mod_2(1) [avatar definition]
340. ~25 | 26 [avatar split clause 194,337,333]
342. 27 <=> m1_subset_1(k6_numbers,k1_enumset1(k6_numbers,1,2)) [avatar definition]
346. 28 <=> k6_numbers = k10_mod_2(k6_numbers) [avatar definition]
348. k6_numbers = k10_mod_2(k6_numbers) <- (28) [avatar component clause 346] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
349. ~27 | 28 [avatar split clause 195,346,342]
351. 29 <=> 1 = k11_mod_2(2,2) [avatar definition]
354. ~23 | 29 [avatar split clause 214,351,324]
356. 30 <=> k6_numbers = k11_mod_2(2,1) [avatar definition]
359. ~23 | ~25 | 30 [avatar split clause 199,356,333,324]
361. 31 <=> k6_numbers = k11_mod_2(1,2) [avatar definition]
364. ~25 | ~23 | 31 [avatar split clause 201,361,324,333]
366. 32 <=> 2 = k11_mod_2(1,1) [avatar definition]
369. ~25 | 32 [avatar split clause 213,366,333]
371. 33 <=> ! [X0] : (k11_mod_2(X0,k6_numbers) = X0 | ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2))) [avatar definition]
372. ~m1_subset_1(X0,k1_enumset1(k6_numbers,1,2)) | k11_mod_2(X0,k6_numbers) = X0 <- (33) [avatar component clause 371] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
373. ~27 | 33 [avatar split clause 204,371,342]
375. 34 <=> ! [X1] : (k11_mod_2(k6_numbers,X1) = X1 | ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2))) [avatar definition]
376. ~m1_subset_1(X1,k1_enumset1(k6_numbers,1,2)) | k11_mod_2(k6_numbers,X1) = X1 <- (34) [avatar component clause 375] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
377. ~27 | 34 [avatar split clause 205,375,342]
389. sK0 != k11_mod_2(sK0,k6_numbers) <- (~2, 4) [forward demodulation 222,231] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
391. 37 <=> sK0 = k11_mod_2(sK0,k6_numbers) [avatar definition]
394. ~37 | 2 | ~4 [avatar split clause 389,229,220,391]
395. m1_subset_1(k6_numbers,k1_enumset1(k6_numbers,1,2)) <- (4, 5) [forward demodulation 236,231] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
396. 27 | ~4 | ~5 [avatar split clause 395,234,229,342]
410. k6_numbers = X0 | ~v1_xboole_0(X0) <- (13) [forward demodulation 150,276] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
416. k10_mod_2(X0) = X0 | ~v1_xboole_0(X0) <- (13, 28) [superposition 348,410] {a:1,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
429. $false <- (~23) [unit resulting resolution 207,326,163] {a:1,w:0,sine:1,thAx:0,allAx:4,thDist:-4}
431. 23 [avatar contradiction clause 429]
433. $false <- (~25) [unit resulting resolution 209,335,163] {a:1,w:0,sine:1,thAx:0,allAx:4,thDist:-4}
435. 25 [avatar contradiction clause 433]
462. v1_xboole_0(k1_enumset1(k6_numbers,1,2)) | r2_hidden(sK0,k1_enumset1(k6_numbers,1,2)) <- (8) [resolution 164,251] {a:1,w:11,wCS:1100,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
463. v1_xboole_0(k1_enumset1(k6_numbers,1,2)) | r2_hidden(sK1,k1_enumset1(k6_numbers,1,2)) <- (7) [resolution 164,246] {a:1,w:11,wCS:1100,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
464. v1_xboole_0(k1_enumset1(k6_numbers,1,2)) | r2_hidden(sK2,k1_enumset1(k6_numbers,1,2)) <- (6) [resolution 164,241] {a:1,w:11,wCS:1100,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
480. r2_hidden(sK0,k1_enumset1(k6_numbers,1,2)) <- (8) [subsumption resolution 462,190] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
482. 42 <=> r2_hidden(sK0,k1_enumset1(k6_numbers,1,2)) [avatar definition]
484. r2_hidden(sK0,k1_enumset1(k6_numbers,1,2)) <- (42) [avatar component clause 482] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
485. 42 | ~8 [avatar split clause 480,249,482]
486. r2_hidden(sK1,k1_enumset1(k6_numbers,1,2)) <- (7) [subsumption resolution 463,190] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
488. 43 <=> r2_hidden(sK1,k1_enumset1(k6_numbers,1,2)) [avatar definition]
490. r2_hidden(sK1,k1_enumset1(k6_numbers,1,2)) <- (43) [avatar component clause 488] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
491. 43 | ~7 [avatar split clause 486,244,488]
492. r2_hidden(sK2,k1_enumset1(k6_numbers,1,2)) <- (6) [subsumption resolution 464,190] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
494. 44 <=> r2_hidden(sK2,k1_enumset1(k6_numbers,1,2)) [avatar definition]
496. r2_hidden(sK2,k1_enumset1(k6_numbers,1,2)) <- (44) [avatar component clause 494] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
497. 44 | ~6 [avatar split clause 492,239,494]
617. 1 = k11_mod_2(1,k6_numbers) <- (25, 33) [resolution 372,334] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:2,thDist:-2}
618. 2 = k11_mod_2(2,k6_numbers) <- (23, 33) [resolution 372,325] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:2,thDist:-2}
619. sK0 = k11_mod_2(sK0,k6_numbers) <- (8, 33) [resolution 372,251] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
620. sK1 = k11_mod_2(sK1,k6_numbers) <- (7, 33) [resolution 372,246] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
637. 37 | ~8 | ~33 [avatar split clause 619,371,249,391]
640. k6_numbers != k11_mod_2(sK0,k10_mod_2(sK0)) <- (~1, 4) [forward demodulation 218,231] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
642. 57 <=> k6_numbers = k11_mod_2(sK0,k10_mod_2(sK0)) [avatar definition]
644. k6_numbers != k11_mod_2(sK0,k10_mod_2(sK0)) <- (~57) [avatar component clause 642] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
645. ~57 | 1 | ~4 [avatar split clause 640,229,216,642]
652. 59 <=> sK1 = k11_mod_2(sK1,k6_numbers) [avatar definition]
662. 59 | ~7 | ~33 [avatar split clause 620,371,244,652]
664. 61 <=> 2 = k11_mod_2(2,k6_numbers) [avatar definition]
667. 61 | ~23 | ~33 [avatar split clause 618,371,324,664]
669. 62 <=> 1 = k11_mod_2(1,k6_numbers) [avatar definition]
672. 62 | ~25 | ~33 [avatar split clause 617,371,333,669]
680. k6_numbers != k11_mod_2(sK0,sK0) | ~v1_xboole_0(sK0) <- (13, 28, ~57) [superposition 644,416] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
682. 64 <=> v1_xboole_0(sK0) [avatar definition]
686. 65 <=> k6_numbers = k11_mod_2(sK0,sK0) [avatar definition]
689. ~64 | ~65 | ~13 | ~28 | 57 [avatar split clause 680,642,346,274,686,682]
691. 1 = k11_mod_2(k6_numbers,1) <- (25, 34) [resolution 376,334] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:2,thDist:-2}
692. 2 = k11_mod_2(k6_numbers,2) <- (23, 34) [resolution 376,325] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:2,thDist:-2}
694. sK1 = k11_mod_2(k6_numbers,sK1) <- (7, 34) [resolution 376,246] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
695. sK2 = k11_mod_2(k6_numbers,sK2) <- (6, 34) [resolution 376,241] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
703. 66 <=> 1 = k11_mod_2(k6_numbers,1) [avatar definition]
706. 66 | ~25 | ~34 [avatar split clause 691,375,333,703]
708. 67 <=> 2 = k11_mod_2(k6_numbers,2) [avatar definition]
711. 67 | ~23 | ~34 [avatar split clause 692,375,324,708]
718. 69 <=> sK1 = k11_mod_2(k6_numbers,sK1) [avatar definition]
721. 69 | ~7 | ~34 [avatar split clause 694,375,244,718]
723. 70 <=> sK2 = k11_mod_2(k6_numbers,sK2) [avatar definition]
726. 70 | ~6 | ~34 [avatar split clause 695,375,239,723]
795. 1 = sK0 | k6_numbers = sK0 | 2 = sK0 <- (42) [resolution 212,484] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
796. 1 = sK1 | k6_numbers = sK1 | 2 = sK1 <- (43) [resolution 212,490] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
797. 1 = sK2 | k6_numbers = sK2 | 2 = sK2 <- (44) [resolution 212,496] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
799. 75 <=> 2 = sK0 [avatar definition]
803. 76 <=> k6_numbers = sK0 [avatar definition]
807. 77 <=> 1 = sK0 [avatar definition]
810. 75 | 76 | 77 | ~42 [avatar split clause 795,482,807,803,799]
812. 78 <=> 2 = sK1 [avatar definition]
816. 79 <=> k6_numbers = sK1 [avatar definition]
820. 80 <=> 1 = sK1 [avatar definition]
823. 78 | 79 | 80 | ~43 [avatar split clause 796,488,820,816,812]
825. 81 <=> 2 = sK2 [avatar definition]
829. 82 <=> k6_numbers = sK2 [avatar definition]
833. 83 <=> 1 = sK2 [avatar definition]
836. 81 | 82 | 83 | ~44 [avatar split clause 797,494,833,829,825]
837. 2 != sK0 | 1 != k10_mod_2(2) | k6_numbers != k11_mod_2(2,1) | k6_numbers = k11_mod_2(sK0,k10_mod_2(sK0)) [theory tautology sat conflict] {a:0,w:18,wCS:7200,sine:255,thAx:0,allAx:1,thDist:-1}
838. 1 != k10_mod_2(2) | 2 != sK0 | 2 != sK1 | 1 != k11_mod_2(2,2) | 2 != sK2 | k6_numbers != k11_mod_2(1,2) | k6_numbers != sK3 | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:43,wCS:17200,sine:255,thAx:0,allAx:1,thDist:-1}
839. 1 != k11_mod_2(2,2) | 1 != sK0 | 2 != sK1 | 2 != k10_mod_2(1) | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | k6_numbers != sK3 | sK2 != k11_mod_2(k6_numbers,sK2) | 2 != sK2 | 2 != k11_mod_2(1,1) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:48,wCS:19200,sine:255,thAx:0,allAx:1,thDist:-1}
840. k6_numbers != k1_xboole_0 | k6_numbers != sK0 | v1_xboole_0(sK0) | ~v1_xboole_0(k1_xboole_0) [theory tautology sat conflict] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:1,thDist:-1}
841. 2 != sK1 | 2 != sK2 | k6_numbers != sK0 | sK1 != k11_mod_2(k6_numbers,sK1) | 1 != k11_mod_2(2,2) | 1 != k11_mod_2(k6_numbers,1) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:35,wCS:14000,sine:255,thAx:0,allAx:1,thDist:-1}
842. k6_numbers != sK3 | sK0 != k11_mod_2(sK0,sK3) | 1 != sK0 | k6_numbers != sK1 | sK2 != k11_mod_2(k6_numbers,sK2) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:30,wCS:12000,sine:255,thAx:0,allAx:1,thDist:-1}
843. 2 != k11_mod_2(1,1) | 1 != sK1 | 2 != sK2 | 2 != k10_mod_2(1) | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | sK0 != k11_mod_2(sK0,sK3) | 1 != sK0 | 1 != k11_mod_2(2,2) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:45,wCS:18000,sine:255,thAx:0,allAx:1,thDist:-1}
844. k6_numbers != sK1 | k6_numbers != sK0 | sK1 != k11_mod_2(k6_numbers,sK1) | sK2 != k11_mod_2(k6_numbers,sK2) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:27,wCS:10800,sine:255,thAx:0,allAx:1,thDist:-1}
845. 1 != sK1 | 2 != sK2 | k6_numbers != sK3 | sK1 != k11_mod_2(k6_numbers,sK1) | sK0 != k11_mod_2(sK0,sK3) | k6_numbers != k11_mod_2(1,2) | k6_numbers != sK0 | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:38,wCS:15200,sine:255,thAx:0,allAx:1,thDist:-1}
846. k6_numbers != sK3 | sK0 != k11_mod_2(sK0,sK3) | 2 != sK0 | k6_numbers != sK1 | sK2 != k11_mod_2(k6_numbers,sK2) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:30,wCS:12000,sine:255,thAx:0,allAx:1,thDist:-1}
847. k6_numbers != k11_mod_2(1,2) | 1 != sK1 | 1 != k10_mod_2(2) | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | k6_numbers != sK3 | sK2 != k11_mod_2(k6_numbers,sK2) | sK0 != k11_mod_2(sK0,sK3) | 2 != sK2 | 2 != sK0 | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:48,wCS:19200,sine:255,thAx:0,allAx:1,thDist:-1}
848. 1 != sK0 | 2 != k10_mod_2(1) | k6_numbers != k11_mod_2(1,2) | k6_numbers = k11_mod_2(sK0,k10_mod_2(sK0)) [theory tautology sat conflict] {a:0,w:18,wCS:7200,sine:255,thAx:0,allAx:1,thDist:-1}
849. k6_numbers != sK3 | sK0 != k11_mod_2(sK0,sK3) | k6_numbers != sK0 | k6_numbers = k11_mod_2(sK0,sK0) [theory tautology sat conflict] {a:0,w:16,wCS:6400,sine:255,thAx:0,allAx:1,thDist:-1}
850. k6_numbers != sK2 | k6_numbers != sK0 | sK1 != k11_mod_2(k6_numbers,sK1) | sK1 != k11_mod_2(sK1,k6_numbers) | 1 != sK1 | 1 != k11_mod_2(k6_numbers,1) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:35,wCS:14000,sine:255,thAx:0,allAx:1,thDist:-1}
851. 1 != sK1 | 1 != sK2 | k6_numbers != sK0 | sK1 != k11_mod_2(k6_numbers,sK1) | 2 != k11_mod_2(1,1) | 2 != k11_mod_2(k6_numbers,2) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:35,wCS:14000,sine:255,thAx:0,allAx:1,thDist:-1}
852. sK1 != k11_mod_2(sK1,k6_numbers) | 2 != sK0 | 1 != sK1 | 1 != k10_mod_2(2) | sK2 != k11_mod_2(k6_numbers,sK2) | k6_numbers != sK2 | k6_numbers != sK3 | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:43,wCS:17200,sine:255,thAx:0,allAx:1,thDist:-1}
853. 2 != k11_mod_2(1,1) | 2 != sK0 | 1 != sK1 | 1 != k10_mod_2(2) | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | k6_numbers != sK3 | sK2 != k11_mod_2(k6_numbers,sK2) | 1 != sK2 | 1 != k11_mod_2(2,2) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:48,wCS:19200,sine:255,thAx:0,allAx:1,thDist:-1}
854. 1 != sK0 | 1 != sK1 | k6_numbers != sK2 | k6_numbers != sK3 | sK0 != k11_mod_2(sK0,sK3) | 2 != k11_mod_2(1,1) | 2 != k11_mod_2(2,k6_numbers) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:38,wCS:15200,sine:255,thAx:0,allAx:1,thDist:-1}
855. 2 != k10_mod_2(1) | 1 != sK1 | 1 != sK0 | 2 != k11_mod_2(1,1) | 1 != sK2 | k6_numbers != k11_mod_2(2,1) | k6_numbers != sK3 | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:43,wCS:17200,sine:255,thAx:0,allAx:1,thDist:-1}
856. sK1 != k11_mod_2(sK1,k6_numbers) | 1 != sK0 | 2 != sK1 | 2 != k10_mod_2(1) | sK2 != k11_mod_2(k6_numbers,sK2) | k6_numbers != sK2 | k6_numbers != sK3 | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:43,wCS:17200,sine:255,thAx:0,allAx:1,thDist:-1}
857. k6_numbers != k11_mod_2(2,1) | 2 != sK1 | 2 != k10_mod_2(1) | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | k6_numbers != sK3 | sK2 != k11_mod_2(k6_numbers,sK2) | sK0 != k11_mod_2(sK0,sK3) | 1 != sK2 | 1 != sK0 | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:48,wCS:19200,sine:255,thAx:0,allAx:1,thDist:-1}
858. 2 != sK0 | 2 != sK1 | k6_numbers != sK2 | k6_numbers != sK3 | sK0 != k11_mod_2(sK0,sK3) | 1 != k11_mod_2(2,2) | 1 != k11_mod_2(1,k6_numbers) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:38,wCS:15200,sine:255,thAx:0,allAx:1,thDist:-1}
859. 1 != k11_mod_2(2,2) | 2 != sK1 | 1 != sK2 | 1 != k10_mod_2(2) | k11_mod_2(sK0,k10_mod_2(sK0)) != sK3 | sK0 != k11_mod_2(sK0,sK3) | 2 != sK0 | 2 != k11_mod_2(1,1) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:45,wCS:18000,sine:255,thAx:0,allAx:1,thDist:-1}
860. k6_numbers != sK2 | k6_numbers != sK0 | sK1 != k11_mod_2(k6_numbers,sK1) | sK1 != k11_mod_2(sK1,k6_numbers) | 2 != sK1 | 2 != k11_mod_2(k6_numbers,2) | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:35,wCS:14000,sine:255,thAx:0,allAx:1,thDist:-1}
861. 2 != sK1 | 1 != sK2 | k6_numbers != sK3 | sK1 != k11_mod_2(k6_numbers,sK1) | sK0 != k11_mod_2(sK0,sK3) | k6_numbers != k11_mod_2(2,1) | k6_numbers != sK0 | k11_mod_2(k11_mod_2(sK0,sK1),sK2) = k11_mod_2(sK0,k11_mod_2(sK1,sK2)) [theory tautology sat conflict] {a:0,w:38,wCS:15200,sine:255,thAx:0,allAx:1,thDist:-1}
862. $false [avatar sat refutation 227,232,237,242,247,252,267,277,331,340,349,354,359,364,369,373,377,394,396,431,435,485,491,497,637,645,662,667,672,689,706,711,721,726,810,823,836,837,838,839,840,841,842,843,844,845,846,847,848,849,850,851,852,853,854,855,856,857,858,859,860,861] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:551,thDist:-551}
% SZS output end Proof for mod_2__l49_mod_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 48
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 70
% Generated clauses: 369
% Active clauses: 92
% Passive clauses: 191
% Final active clauses: 86
% Final passive clauses: 86
% Main loop iterations started: 101
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Fw subsumption resolutions: 33
% Fw demodulations: 7
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Equational tautologies: 3
% Forward subsumptions: 30
% 
% >>> Generating Inferences
% Binary resolution: 121
% Unit resulting resolution: 3
% Forward superposition: 27
% Backward superposition: 10
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 27
% Split components: 61
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 116
% SAT solver unit clauses: 23
% SAT solver binary clauses: 16
% 
% Memory used [KB]: 5500
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% congruence closure: 0.001 s
% forward subsumption: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

