i: [1,34,0,2,1,1,0,0] t34_xxreal_0
a: [3,53,0,2,1,34]
s: 53
i: [1,33,0,2,1,1,0,0] t34_xxreal_0
a: [3,58,0,2,1,33]
s: 58
i: [1,32,0,2,1,1,0,0] t34_xxreal_0
a: [3,63,0,2,1,32]
s: 63
i: [1,44,0,2,1,0,0,1] rc1_xxreal_0
a: [3,79,0,2,1,44]
s: 79
i: [1,35,0,11,1,1,0,0] t34_xxreal_0
a: [3,48,0,11,1,35]
s: 48
i: [1,39,0,7,3,0,0,2] reflexivity_r1_xxreal_0
a: [3,70,0,5,2,39]
s: 70
i: [1,41,0,9,3,0,0,1] idempotence_k4_xxreal_0
a: [3,74,0,7,2,41]
s: 74
i: [1,40,0,8,3,0,0,1] fc6_xxreal_0
s: 40
i: [1,43,0,10,4,0,0,2] connectedness_r1_xxreal_0
s: 43
i: [1,42,0,11,3,0,0,1] commutativity_k4_xxreal_0
s: 42
i: [1,36,0,12,4,0,0,1] d10_xxreal_0
s: 36
d: [2,105,1,16,4,1,44,48,36]
d: [2,115,1,14,3,2,21,105,58]
d: [2,116,1,12,2,3,21,115,63]
a: [3,120,1,3,1,116]
s: 120
r: 120
i: [1,37,0,12,4,0,0,1] d10_xxreal_0
s: 37
a: [3,119,1,3,1,116]
s: 119
i: [1,38,0,15,6,0,0,2] t2_xxreal_0
s: 38
d: [2,85,1,17,3,1,44,48,42]
d: [2,91,1,15,2,2,21,85,53]
a: [3,95,1,4,1,91]
s: 95
d: [2,180,2,0,0,3,52,63,58,95,40]
e: 180
d: [2,181,2,4,2,1,40,95,40]
d: [2,187,2,2,1,2,21,181,58]
d: [2,188,2,0,0,3,21,187,63]
e: 188
d: [2,183,2,9,4,1,44,95,36]
d: [2,184,2,7,3,1,18,183]
d: [2,190,2,5,2,2,21,184,63]
d: [2,191,2,2,1,3,21,190,119]
d: [2,192,2,0,0,4,21,191,58]
e: 192
r: 95
a: [3,94,1,4,1,91]
s: 94
d: [2,106,1,20,4,1,44,48,36]
d: [2,126,1,18,3,2,21,106,53]
a: [3,130,1,5,1,126]
s: 130
d: [2,198,2,10,4,1,44,130,36]
d: [2,206,2,7,3,2,21,198,119]
d: [2,207,2,5,2,3,21,206,58]
d: [2,208,2,3,1,4,21,207,63]
a: [3,212,2,3,1,208]
s: 212
r: 119
r: 212
s: 120
d: [2,148,1,18,4,1,44,48,37]
d: [2,167,1,16,3,2,21,148,53]
a: [3,171,1,7,1,167]
s: 171
d: [2,136,2,7,3,1,40,120,43]
d: [2,137,2,5,2,2,21,136,58]
d: [2,138,2,3,1,3,21,137,63]
a: [3,142,2,3,1,138]
s: 142
a: [3,124,1,9,1,116]
s: 124
r: 130
a: [3,211,2,3,1,208]
s: 211
a: [3,129,1,5,1,126]
s: 129
a: [3,134,1,9,1,126]
s: 134
r: 120
r: 129
r: 134
r: 211
s: 119
s: 130
s: 212
a: [3,133,1,9,1,126]
s: 133
r: 124
r: 48
d: [2,357,0,9,1,2,23,48,133]
a: [3,369,0,9,1,357]
s: 369
d: [2,197,2,10,4,1,44,130,37]
d: [2,225,2,8,3,2,21,197,58]
d: [2,226,2,6,2,3,21,225,63]
a: [3,229,2,3,1,226]
s: 229
d: [2,408,3,0,0,6,52,53,58,63,119,212,229,38]
e: 408
r: 119
r: 229
s: 120
a: [3,230,2,3,1,226]
s: 230
d: [2,383,1,12,4,1,44,369,37]
d: [2,416,1,10,3,2,21,383,58]
d: [2,417,1,8,2,3,21,416,63]
a: [3,421,1,5,1,417]
s: 421
d: [2,147,1,16,4,1,44,48,37]
d: [2,160,1,14,3,2,21,147,58]
d: [2,161,1,12,2,3,21,160,63]
a: [3,165,1,9,1,161]
s: 165
d: [2,437,2,10,4,1,44,421,37]
d: [2,440,2,7,3,2,21,437,230]
d: [2,441,2,5,2,3,21,440,53]
d: [2,442,2,3,1,4,21,441,58]
a: [3,446,2,3,1,442]
s: 446
d: [2,387,1,15,3,1,44,369,42]
d: [2,398,1,11,2,2,21,387,94]
d: [2,399,1,9,1,3,21,398,53]
a: [3,403,1,9,1,399]
s: 403
d: [2,214,3,7,3,1,40,212,43]
d: [2,215,3,5,2,2,21,214,53]
d: [2,216,3,3,1,3,21,215,63]
a: [3,220,3,3,1,216]
s: 220
a: [3,99,1,11,1,91]
s: 99
d: [2,108,1,16,6,0,44,42,36]
d: [2,111,1,12,4,0,18,108]
s: 111
d: [2,429,3,7,3,1,40,230,43]
d: [2,430,3,5,2,2,21,429,53]
d: [2,431,3,3,1,3,21,430,58]
a: [3,435,3,3,1,431]
s: 435
d: [2,107,1,16,6,0,44,42,36]
d: [2,112,1,12,4,0,18,107]
s: 112
d: [2,238,2,16,4,1,44,171,37]
d: [2,254,2,14,3,2,21,238,63]
a: [3,258,2,4,1,254]
s: 258
d: [2,526,3,0,0,3,52,58,53,258,40]
e: 526
d: [2,527,3,4,2,1,40,258,40]
d: [2,535,3,2,1,2,21,527,53]
d: [2,536,3,0,0,3,21,535,58]
e: 536
d: [2,528,3,9,4,1,44,258,111]
d: [2,533,3,7,3,1,18,528]
d: [2,538,3,5,2,2,21,533,53]
d: [2,539,3,3,1,3,21,538,58]
d: [2,540,3,0,0,4,21,539,435]
e: 540
d: [2,529,3,9,4,1,44,258,37]
d: [2,532,3,7,3,1,18,529]
d: [2,542,3,5,2,2,21,532,53]
d: [2,543,3,2,1,3,21,542,230]
d: [2,544,3,0,0,4,21,543,58]
e: 544
r: 133
r: 258
s: 48
a: [3,123,1,9,1,116]
s: 123
r: 48
d: [2,239,2,14,4,1,44,171,36]
d: [2,268,2,12,3,2,21,239,63]
a: [3,272,2,3,1,268]
s: 272
d: [2,569,0,9,1,2,23,48,123]
a: [3,582,0,9,1,569]
s: 582
a: [3,257,2,4,1,254]
s: 257
s: 134
s: 165
d: [2,578,2,10,3,1,44,40,123]
d: [2,597,2,8,2,2,21,578,63]
a: [3,601,2,4,1,597]
s: 601
d: [2,609,1,15,3,1,44,582,42]
d: [2,627,1,11,2,2,21,609,94]
d: [2,628,1,9,1,3,21,627,53]
a: [3,632,1,9,1,628]
s: 632
d: [2,196,2,11,3,1,40,130,43]
d: [2,199,2,9,2,2,21,196,53]
d: [2,200,2,5,1,3,21,199,94]
a: [3,204,2,5,1,200]
s: 204
s: 99
d: [2,104,1,16,6,0,44,36,42]
d: [2,113,1,12,4,0,18,104]
s: 113
d: [2,236,2,12,4,1,44,171,37]
d: [2,240,2,10,3,2,21,236,53]
d: [2,241,2,8,2,3,21,240,58]
a: [3,245,2,5,1,241]
s: 245
d: [2,702,4,0,0,4,52,63,53,220,245,111]
e: 702
d: [2,703,3,0,0,4,52,63,53,212,245,37]
e: 703
d: [2,704,3,10,4,1,44,245,111]
d: [2,708,3,7,3,1,19,704]
d: [2,711,3,5,2,2,21,708,63]
d: [2,712,3,3,1,3,21,711,53]
d: [2,713,3,0,0,4,21,712,220]
e: 713
d: [2,705,3,10,4,1,44,245,37]
d: [2,707,3,7,3,1,19,705]
d: [2,715,3,4,2,2,21,707,212]
d: [2,716,3,2,1,3,21,715,53]
d: [2,717,3,0,0,4,21,716,63]
e: 717
r: 212
r: 220
d: [2,103,1,16,6,0,44,36,42]
d: [2,114,1,12,4,0,18,103]
s: 114
s: 211
d: [2,723,3,0,0,6,52,53,63,58,142,230,211,38]
e: 723
d: [2,724,4,0,0,6,52,58,53,63,435,120,211,38]
e: 724
d: [2,277,2,14,4,1,44,124,37]
d: [2,281,2,7,3,1,19,277]
d: [2,282,2,5,2,2,21,281,53]
d: [2,283,2,3,1,3,21,282,58]
d: [2,308,2,10,4,1,44,129,37]
d: [2,312,2,7,3,2,21,308,120]
d: [2,313,2,5,2,3,21,312,58]
d: [2,314,2,3,1,4,21,313,63]
d: [2,317,2,14,4,1,44,134,36]
d: [2,320,2,7,3,1,19,317]
d: [2,324,2,5,2,2,21,320,53]
d: [2,325,2,3,1,3,21,324,58]
d: [2,344,2,10,4,1,44,130,36]
d: [2,348,2,7,3,2,21,344,119]
d: [2,349,2,5,2,3,21,348,58]
d: [2,350,2,3,1,4,21,349,63]
d: [2,343,2,10,4,1,44,130,37]
d: [2,423,2,8,3,2,21,343,58]
d: [2,424,2,6,2,3,21,423,63]
d: [2,607,1,16,4,1,44,582,37]
d: [2,623,1,11,3,2,21,607,130]
d: [2,624,1,9,2,3,21,623,53]
d: [2,625,1,5,1,4,21,624,94]
d: [2,719,3,5,2,2,21,707,53]
d: [2,720,3,3,1,3,21,719,63]
f: 724,720,625,569,431,424,408,350,325,314,283,188,126,116,32,33,34,35
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_0__t34_xxreal_0
% SZS output start Proof for xxreal_0__t34_xxreal_0
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => k4_xxreal_0(k4_xxreal_0(X0,X1),X2) = k4_xxreal_0(X0,k4_xxreal_0(X1,X2))))) [input t34_xxreal_0]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => k4_xxreal_0(k4_xxreal_0(X0,X1),X2) = k4_xxreal_0(X0,k4_xxreal_0(X1,X2))))) [negated conjecture 1]
4. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
5. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((~r1_xxreal_0(X1,X0) => k4_xxreal_0(X0,X1) = X1) & (r1_xxreal_0(X1,X0) => k4_xxreal_0(X0,X1) = X0)))) [input d10_xxreal_0]
7. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v1_xxreal_0(k4_xxreal_0(X0,X1))) [input fc6_xxreal_0]
11. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input t2_xxreal_0]
12. ? [X0] : (? [X1] : (? [X2] : (k4_xxreal_0(k4_xxreal_0(X0,X1),X2) != k4_xxreal_0(X0,k4_xxreal_0(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
13. ! [X0] : (! [X1] : (((k4_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X1,X0)) & (k4_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X1,X0))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 5]
14. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 11]
15. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 14]
18. ! [X0,X1] : (v1_xxreal_0(k4_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 7]
19. ! [X0,X1] : (v1_xxreal_0(k4_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 18]
24. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 4]
25. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 24]
26. ? [X0] : (? [X1] : (? [X2] : (k4_xxreal_0(k4_xxreal_0(X0,X1),X2) != k4_xxreal_0(X0,k4_xxreal_0(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (k4_xxreal_0(k4_xxreal_0(sK0,X1),X2) != k4_xxreal_0(sK0,k4_xxreal_0(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
27. ? [X1] : (? [X2] : (k4_xxreal_0(k4_xxreal_0(sK0,X1),X2) != k4_xxreal_0(sK0,k4_xxreal_0(X1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (k4_xxreal_0(k4_xxreal_0(sK0,sK1),X2) != k4_xxreal_0(sK0,k4_xxreal_0(sK1,X2)) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
28. ? [X2] : (k4_xxreal_0(k4_xxreal_0(sK0,sK1),X2) != k4_xxreal_0(sK0,k4_xxreal_0(sK1,X2)) & v1_xxreal_0(X2)) => (k4_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) != k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) & v1_xxreal_0(sK2)) [choice axiom]
29. ((k4_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) != k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 12,28,27,26]
32. v1_xxreal_0(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
33. v1_xxreal_0(sK1) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
34. v1_xxreal_0(sK2) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
35. k4_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) != k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) [cnf transformation 29] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
36. k4_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 13] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. k4_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 13] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
38. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 15] {a:0,w:15,wCS:6000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
40. v1_xxreal_0(k4_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 19] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 25] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
46. 1 <=> k4_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) = k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) [avatar definition]
48. k4_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) != k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) <- (~1) [avatar component clause 46] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
49. ~1 [avatar split clause 35,46]
51. 2 <=> v1_xxreal_0(sK2) [avatar definition]
53. v1_xxreal_0(sK2) <- (2) [avatar component clause 51] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
54. 2 [avatar split clause 34,51]
56. 3 <=> v1_xxreal_0(sK1) [avatar definition]
58. v1_xxreal_0(sK1) <- (3) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
59. 3 [avatar split clause 33,56]
61. 4 <=> v1_xxreal_0(sK0) [avatar definition]
63. v1_xxreal_0(sK0) <- (4) [avatar component clause 61] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
64. 4 [avatar split clause 32,61]
93. 9 <=> v1_xxreal_0(k4_xxreal_0(sK0,sK1)) [avatar definition]
94. v1_xxreal_0(k4_xxreal_0(sK0,sK1)) <- (9) [avatar component clause 93] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
95. ~v1_xxreal_0(k4_xxreal_0(sK0,sK1)) <- (~9) [avatar component clause 93] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
105. k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) != k4_xxreal_0(sK0,sK2) | ~r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~1) [superposition 48,36] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
106. k4_xxreal_0(sK0,sK1) != k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) | ~r1_xxreal_0(sK2,k4_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(k4_xxreal_0(sK0,sK1)) <- (~1) [superposition 48,36] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
115. k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) != k4_xxreal_0(sK0,sK2) | ~r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) <- (~1, 3) [subsumption resolution 105,58] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
116. k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) != k4_xxreal_0(sK0,sK2) | ~r1_xxreal_0(sK1,sK0) <- (~1, 3, 4) [subsumption resolution 115,63] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
118. 11 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
119. r1_xxreal_0(sK1,sK0) <- (11) [avatar component clause 118] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
120. ~r1_xxreal_0(sK1,sK0) <- (~11) [avatar component clause 118] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
122. 12 <=> k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) = k4_xxreal_0(sK0,sK2) [avatar definition]
123. k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) = k4_xxreal_0(sK0,sK2) <- (12) [avatar component clause 122] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
124. k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) != k4_xxreal_0(sK0,sK2) <- (~12) [avatar component clause 122] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
125. ~11 | ~12 | 1 | ~3 | ~4 [avatar split clause 116,61,56,46,122,118]
126. k4_xxreal_0(sK0,sK1) != k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) | ~r1_xxreal_0(sK2,k4_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(k4_xxreal_0(sK0,sK1)) <- (~1, 2) [subsumption resolution 106,53] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
128. 13 <=> r1_xxreal_0(sK2,k4_xxreal_0(sK0,sK1)) [avatar definition]
129. r1_xxreal_0(sK2,k4_xxreal_0(sK0,sK1)) <- (13) [avatar component clause 128] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
130. ~r1_xxreal_0(sK2,k4_xxreal_0(sK0,sK1)) <- (~13) [avatar component clause 128] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
132. 14 <=> k4_xxreal_0(sK0,sK1) = k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) [avatar definition]
134. k4_xxreal_0(sK0,sK1) != k4_xxreal_0(sK0,k4_xxreal_0(sK1,sK2)) <- (~14) [avatar component clause 132] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
135. ~9 | ~13 | ~14 | 1 | ~2 [avatar split clause 126,51,46,132,128,93]
181. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~9) [resolution 95,40] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
187. ~v1_xxreal_0(sK0) <- (3, ~9) [subsumption resolution 181,58] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
188. $false <- (3, 4, ~9) [subsumption resolution 187,63] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
189. ~3 | ~4 | 9 [avatar contradiction clause 188]
210. 19 <=> r1_xxreal_0(sK2,sK0) [avatar definition]
211. r1_xxreal_0(sK2,sK0) <- (19) [avatar component clause 210] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
212. ~r1_xxreal_0(sK2,sK0) <- (~19) [avatar component clause 210] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
228. 21 <=> r1_xxreal_0(sK2,sK1) [avatar definition]
229. r1_xxreal_0(sK2,sK1) <- (21) [avatar component clause 228] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
230. ~r1_xxreal_0(sK2,sK1) <- (~21) [avatar component clause 228] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
243. 22 <=> sK2 = k4_xxreal_0(sK0,sK2) [avatar definition]
245. sK2 != k4_xxreal_0(sK0,sK2) <- (~22) [avatar component clause 243] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
277. k4_xxreal_0(sK0,sK2) != k4_xxreal_0(sK0,sK2) | r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~12) [superposition 124,37] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
281. r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~12) [trivial inequality removal 277] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
282. r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK1) <- (2, ~12) [subsumption resolution 281,53] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
283. r1_xxreal_0(sK2,sK1) <- (2, 3, ~12) [subsumption resolution 282,58] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
284. 21 | ~2 | ~3 | 12 [avatar split clause 283,122,56,51,228]
308. r1_xxreal_0(sK2,sK1) | r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (13) [superposition 129,37] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
312. r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~11, 13) [subsumption resolution 308,120] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
313. r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK0) <- (3, ~11, 13) [subsumption resolution 312,58] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
314. r1_xxreal_0(sK2,sK1) <- (3, 4, ~11, 13) [subsumption resolution 313,63] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
315. 21 | ~3 | ~4 | 11 | ~13 [avatar split clause 314,128,118,61,56,228]
317. k4_xxreal_0(sK0,sK1) != k4_xxreal_0(sK0,sK1) | ~r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~14) [superposition 134,36] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
320. ~r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~14) [trivial inequality removal 317] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
324. ~r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK1) <- (2, ~14) [subsumption resolution 320,53] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
325. ~r1_xxreal_0(sK2,sK1) <- (2, 3, ~14) [subsumption resolution 324,58] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
326. ~21 | ~2 | ~3 | 14 [avatar split clause 325,132,56,51,228]
343. ~r1_xxreal_0(sK2,sK1) | r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~13) [superposition 130,37] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
344. ~r1_xxreal_0(sK2,sK0) | ~r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~13) [superposition 130,36] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
348. ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (11, ~13) [subsumption resolution 344,119] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
349. ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) <- (3, 11, ~13) [subsumption resolution 348,58] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
350. ~r1_xxreal_0(sK2,sK0) <- (3, 4, 11, ~13) [subsumption resolution 349,63] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
351. ~19 | ~3 | ~4 | ~11 | 13 [avatar split clause 350,128,118,61,56,210]
408. $false <- (2, 3, 4, 11, ~19, 21) [unit resulting resolution 53,58,63,119,212,229,38] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:30,thDist:-30}
410. ~2 | ~3 | ~4 | ~11 | 19 | ~21 [avatar contradiction clause 408]
423. ~r1_xxreal_0(sK2,sK1) | r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) <- (3, ~13) [subsumption resolution 343,58] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
424. ~r1_xxreal_0(sK2,sK1) | r1_xxreal_0(sK1,sK0) <- (3, 4, ~13) [subsumption resolution 423,63] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
425. 11 | ~21 | ~3 | ~4 | 13 [avatar split clause 424,128,61,56,228,118]
429. r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~21) [resolution 230,43] {a:3,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
430. r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK1) <- (2, ~21) [subsumption resolution 429,53] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
431. r1_xxreal_0(sK1,sK2) <- (2, 3, ~21) [subsumption resolution 430,58] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
433. 36 <=> r1_xxreal_0(sK1,sK2) [avatar definition]
435. r1_xxreal_0(sK1,sK2) <- (36) [avatar component clause 433] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
436. 36 | ~2 | ~3 | 21 [avatar split clause 431,228,56,51,433]
569. k4_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) != k4_xxreal_0(sK0,sK2) <- (~1, 12) [backward demodulation 48,123] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
580. 40 <=> k4_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) = k4_xxreal_0(sK0,sK2) [avatar definition]
582. k4_xxreal_0(k4_xxreal_0(sK0,sK1),sK2) != k4_xxreal_0(sK0,sK2) <- (~40) [avatar component clause 580] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
583. ~40 | 1 | ~12 [avatar split clause 569,122,46,580]
607. sK2 != k4_xxreal_0(sK0,sK2) | r1_xxreal_0(sK2,k4_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(k4_xxreal_0(sK0,sK1)) <- (~40) [superposition 582,37] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
623. sK2 != k4_xxreal_0(sK0,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(k4_xxreal_0(sK0,sK1)) <- (~13, ~40) [subsumption resolution 607,130] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
624. sK2 != k4_xxreal_0(sK0,sK2) | ~v1_xxreal_0(k4_xxreal_0(sK0,sK1)) <- (2, ~13, ~40) [subsumption resolution 623,53] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
625. sK2 != k4_xxreal_0(sK0,sK2) <- (2, 9, ~13, ~40) [subsumption resolution 624,94] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
626. ~22 | ~2 | ~9 | 13 | 40 [avatar split clause 625,580,128,93,51,243]
705. sK2 != sK2 | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (~22) [superposition 245,37] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
707. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (~22) [trivial inequality removal 705] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
719. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) <- (2, ~22) [subsumption resolution 707,53] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
720. r1_xxreal_0(sK2,sK0) <- (2, 4, ~22) [subsumption resolution 719,63] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
721. 19 | ~2 | ~4 | 22 [avatar split clause 720,243,61,51,210]
724. $false <- (2, 3, 4, ~11, 19, 36) [unit resulting resolution 58,53,63,435,120,211,38] {a:4,w:0,goal:1,sine:2,thAx:0,allAx:31,thDist:-31}
727. ~2 | ~3 | ~4 | 11 | ~19 | ~36 [avatar contradiction clause 724]
730. $false [avatar sat refutation 49,54,59,64,125,135,189,284,315,326,351,410,425,436,583,626,721,727] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1197,thDist:-1197}
% SZS output end Proof for xxreal_0__t34_xxreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 492
% Active clauses: 52
% Passive clauses: 89
% Final active clauses: 34
% Final passive clauses: 21
% Main loop iterations started: 59
% 
% >>> Simplifying Inferences
% Duplicate literals: 54
% Trivial inequalities: 10
% Fw subsumption resolutions: 192
% Fw demodulations: 5
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 16
% Forward subsumptions: 44
% 
% >>> Generating Inferences
% Binary resolution: 22
% Unit resulting resolution: 7
% Factoring: 1
% Forward superposition: 95
% Backward superposition: 39
% 
% >>> AVATAR
% Split clauses: 48
% Split components: 117
% Sat splitting refutations: 14
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 114
% SAT solver unit clauses: 5
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 5245
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% term sharing: 0.001 s
% forward subsumption: 0.001 s
% forward demodulation: 0.001 s
% superposition: 0.001 s ( own 0.0000 s ) 
% other: 0.005 s ( own 0.001 s ) 

