i: [1,128,0,2,1,1,0,0] t23_ordinal3
a: [3,195,0,2,1,128]
s: 195
i: [1,127,0,2,1,1,0,0] t23_ordinal3
a: [3,200,0,2,1,127]
s: 200
i: [1,126,0,2,1,1,0,0] t23_ordinal3
a: [3,205,0,2,1,126]
s: 205
i: [1,130,0,3,1,1,0,0] t23_ordinal3
a: [3,185,0,3,1,130]
s: 185
i: [1,129,0,7,1,1,0,0] t23_ordinal3
a: [3,190,0,7,1,129]
s: 190
i: [1,131,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,210,0,2,1,131]
s: 210
i: [1,132,0,2,1,0,0,255] fc1_xboole_0
a: [3,215,0,2,1,132]
s: 215
i: [1,172,0,2,1,0,0,3] rc2_xboole_0
a: [3,232,0,2,1,172]
s: 232
i: [1,173,0,2,1,0,0,2] rc4_ordinal1
a: [3,237,0,2,1,173]
s: 237
i: [1,174,0,2,1,0,0,3] rc1_xboole_0
a: [3,242,0,2,1,174]
s: 242
i: [1,175,0,2,1,0,0,1] rc1_ordinal1
a: [3,247,0,2,1,175]
s: 247
i: [1,179,0,2,1,0,0,255] rc2_ordinal1
a: [3,252,0,2,1,179]
s: 252
i: [1,178,0,2,1,0,0,255] rc2_ordinal1
a: [3,257,0,2,1,178]
s: 257
i: [1,177,0,2,1,0,0,255] rc2_ordinal1
a: [3,262,0,2,1,177]
s: 262
i: [1,176,0,2,1,0,0,255] rc2_ordinal1
a: [3,267,0,2,1,176]
s: 267
i: [1,181,0,2,1,0,0,2] rc5_ordinal1
a: [3,272,0,2,1,181]
s: 272
i: [1,180,0,2,1,0,0,2] rc5_ordinal1
a: [3,277,0,2,1,180]
s: 277
i: [1,161,0,3,1,0,0,255] irreflexivity_r2_xboole_0
s: 161
i: [1,162,0,3,1,0,0,2] reflexivity_r1_tarski
s: 162
i: [1,133,0,3,1,0,0,255] d2_xboole_0
a: [3,220,0,3,1,133]
s: 220
r: 210
i: [1,136,0,4,2,0,0,2] cc6_ordinal1
s: 136
d: [2,280,1,2,1,1,40,136,237]
a: [3,285,1,2,1,280]
s: 285
i: [1,137,0,4,2,0,0,255] cc9_ordinal1
s: 137
d: [2,281,1,2,1,1,40,136,272]
a: [3,290,1,2,1,281]
s: 290
i: [1,138,0,4,2,0,0,2] cc7_ordinal1
s: 138
i: [1,139,0,4,2,0,0,3] cc3_ordinal1
s: 139
d: [2,292,1,2,1,1,40,137,242]
a: [3,297,1,2,1,292]
s: 297
i: [1,141,0,4,2,0,0,255] cc1_ordinal1
s: 141
d: [2,328,1,2,1,1,40,141,205]
a: [3,338,1,2,1,328]
s: 338
i: [1,142,0,4,2,0,0,255] cc1_ordinal1
s: 142
i: [1,160,0,4,1,0,0,3] existence_m1_subset_1
s: 160
d: [2,329,1,2,1,1,40,141,200]
a: [3,343,1,2,1,329]
s: 343
i: [1,140,0,5,2,0,0,255] t6_boole
s: 140
d: [2,330,1,2,1,1,40,141,195]
a: [3,348,1,2,1,330]
s: 348
i: [1,157,0,5,2,0,0,2] t7_boole
s: 157
i: [1,167,0,7,3,0,0,1] reflexivity_r1_ordinal1
a: [3,227,0,5,2,167]
s: 227
d: [2,365,1,2,1,1,40,142,205]
a: [3,375,1,2,1,365]
s: 375
i: [1,147,0,6,3,0,0,255] cc2_ordinal1
s: 147
d: [2,366,1,2,1,1,40,142,200]
a: [3,380,1,2,1,366]
s: 380
i: [1,149,0,6,2,0,0,2] d2_ordinal1
s: 149
i: [1,150,0,6,2,0,0,2] d2_ordinal1
s: 150
d: [2,367,1,2,1,1,40,142,195]
a: [3,385,1,2,1,367]
s: 385
i: [1,152,0,6,2,0,0,2] t1_subset
s: 152
d: [2,293,1,2,1,1,40,137,215]
a: [3,302,1,2,1,293]
s: 302
i: [1,163,0,6,2,0,0,255] antisymmetry_r2_xboole_0
s: 163
i: [1,164,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 164
d: [2,304,1,2,1,1,40,138,242]
a: [3,309,1,2,1,304]
s: 309
i: [1,145,0,7,3,0,0,3] cc2_ordinal2
s: 145
d: [2,305,1,2,1,1,40,138,215]
a: [3,314,1,2,1,305]
s: 314
i: [1,146,0,7,3,0,0,3] cc5_ordinal1
s: 146
i: [1,154,0,7,2,0,0,255] t3_subset
s: 154
d: [2,316,1,2,1,1,40,139,242]
a: [3,321,1,2,1,316]
s: 321
i: [1,155,0,7,2,0,0,255] t3_subset
s: 155
d: [2,317,1,2,1,1,40,139,215]
a: [3,326,1,2,1,317]
s: 326
i: [1,156,0,7,3,0,0,2] t8_boole
s: 156
i: [1,134,0,8,3,0,0,255] cc10_ordinal1
s: 134
d: [2,332,1,2,1,1,40,141,247]
a: [3,358,1,2,1,332]
s: 358
i: [1,148,0,8,3,0,0,2] d2_ordinal1
s: 148
d: [2,369,1,2,1,1,40,142,247]
a: [3,395,1,2,1,369]
s: 395
i: [1,151,0,8,3,0,0,2] t2_subset
s: 151
i: [1,165,0,8,3,0,0,1] fc5_ordinal2
s: 165
d: [2,331,2,2,1,1,40,141,285]
a: [3,353,2,2,1,331]
s: 353
i: [1,166,0,8,3,0,0,1] fc5_ordinal2
s: 166
d: [2,334,2,2,1,1,40,141,290]
a: [3,363,2,2,1,334]
s: 363
i: [1,168,0,8,3,0,0,1] dt_k10_ordinal2
s: 168
i: [1,153,0,9,3,0,0,255] d8_xboole_0
s: 153
d: [2,368,2,2,1,1,40,142,285]
a: [3,390,2,2,1,368]
s: 390
i: [1,159,0,9,3,0,0,255] t5_subset
s: 159
d: [2,371,2,2,1,1,40,142,290]
a: [3,400,2,2,1,371]
s: 400
i: [1,135,0,10,4,0,0,255] t11_ordinal1
s: 135
i: [1,158,0,10,3,0,0,255] t4_subset
s: 158
d: [2,419,2,2,1,1,40,321,142]
a: [3,424,2,2,1,419]
s: 424
i: [1,169,0,10,4,0,0,1] connectedness_r1_ordinal1
s: 169
d: [2,465,1,7,3,1,40,169,185]
d: [2,468,1,5,2,2,21,465,200]
d: [2,469,1,3,1,3,21,468,195]
a: [3,473,1,3,1,469]
s: 473
i: [1,170,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 170
i: [1,171,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 171
d: [2,478,2,7,3,1,40,170,473]
d: [2,494,2,5,2,2,21,478,200]
d: [2,495,2,3,1,3,21,494,195]
a: [3,499,2,3,1,495]
s: 499
i: [1,143,0,16,5,0,0,1] t21_ordinal3
s: 143
d: [2,504,3,6,2,1,40,499,153]
a: [3,508,3,3,1,504]
s: 508
i: [1,144,0,16,5,0,0,1] t22_ordinal3
s: 144
d: [2,475,1,15,3,1,40,170,190]
a: [3,484,1,4,1,475]
s: 484
d: [2,531,2,0,0,3,52,205,200,484,168]
e: 531
d: [2,532,2,4,2,1,40,484,168]
d: [2,535,2,2,1,2,21,532,200]
d: [2,536,2,0,0,3,21,535,205]
e: 536
r: 484
a: [3,511,3,3,1,504]
s: 511
a: [3,483,1,4,1,475]
s: 483
d: [2,516,4,7,3,1,40,508,135]
d: [2,518,4,5,2,2,21,516,200]
d: [2,519,4,3,1,3,21,518,348]
a: [3,523,4,3,1,519]
s: 523
a: [3,488,1,4,1,475]
s: 488
d: [2,569,2,0,0,3,52,205,195,488,168]
e: 569
d: [2,570,2,4,2,1,40,488,168]
d: [2,573,2,2,1,2,21,570,195]
d: [2,574,2,0,0,3,21,573,205]
e: 574
r: 488
a: [3,487,1,4,1,475]
s: 487
d: [2,553,5,2,1,1,40,523,157]
a: [3,567,5,2,1,553]
s: 567
a: [3,492,1,7,1,475]
s: 492
d: [2,517,4,3,1,1,40,508,163]
a: [3,528,4,3,1,517]
s: 528
d: [2,407,1,4,2,0,40,149,157]
s: 407
d: [2,414,1,5,2,0,40,145,160]
s: 414
d: [2,551,5,3,1,1,40,523,164]
a: [3,557,5,3,1,551]
s: 557
d: [2,418,1,5,1,0,40,154,160]
s: 418
d: [2,552,5,3,1,1,40,523,152]
a: [3,562,5,3,1,552]
s: 562
d: [2,409,1,6,2,0,40,152,149]
s: 409
d: [2,411,1,6,2,0,40,164,149]
s: 411
d: [2,538,2,4,1,1,40,483,142]
a: [3,543,2,4,1,538]
s: 543
d: [2,446,1,6,2,0,40,134,160]
s: 446
d: [2,539,2,4,1,1,40,483,141]
a: [3,548,2,4,1,539]
s: 548
d: [2,450,1,6,2,0,40,151,160]
s: 450
d: [2,402,1,7,3,0,44,140,140]
s: 402
d: [2,576,2,4,1,1,40,487,142]
a: [3,581,2,4,1,576]
s: 581
d: [2,447,1,7,3,0,40,134,155]
s: 447
d: [2,623,3,4,2,1,40,447,499]
a: [3,641,3,2,1,623]
s: 641
d: [2,459,1,7,2,0,40,159,160]
s: 459
d: [2,432,1,8,3,0,40,155,145]
s: 432
d: [2,577,2,4,1,1,40,487,141]
a: [3,586,2,4,1,577]
s: 586
d: [2,452,1,8,3,0,40,165,142]
s: 452
d: [2,622,2,8,2,1,40,447,492]
a: [3,632,2,4,1,622]
s: 632
d: [2,453,1,8,3,0,40,165,141]
s: 453
d: [2,455,1,8,3,0,40,168,142]
s: 455
d: [2,589,2,14,2,1,40,492,153]
a: [3,593,2,7,1,589]
s: 593
d: [2,456,1,8,3,0,40,168,141]
s: 456
d: [2,651,3,15,3,1,40,593,135]
d: [2,653,3,11,2,2,21,651,487]
d: [2,654,3,7,1,3,21,653,548]
a: [3,658,3,7,1,654]
s: 658
d: [2,665,6,0,0,5,52,205,200,195,557,658,144]
e: 665
d: [2,666,4,9,4,1,40,658,144]
d: [2,672,4,6,3,2,21,666,557]
d: [2,673,4,4,2,3,21,672,195]
d: [2,674,4,2,1,4,21,673,200]
d: [2,675,4,0,0,5,21,674,205]
e: 675
r: 593
r: 658
d: [2,460,1,8,3,0,40,159,155]
s: 460
d: [2,462,1,8,2,0,40,158,160]
s: 462
d: [2,678,2,9,2,1,40,460,492]
a: [3,684,2,4,1,678]
s: 684
d: [2,463,1,9,3,0,40,158,155]
s: 463
d: [2,693,3,6,2,1,40,463,499]
s: 693
d: [2,451,1,10,3,0,40,151,155]
s: 451
d: [2,477,1,14,6,0,40,170,169]
d: [2,479,1,10,4,0,18,477]
s: 479
d: [2,696,4,6,2,1,40,693,450]
a: [3,700,4,2,1,696]
s: 700
r: 700
a: [3,592,2,7,1,589]
s: 592
a: [3,699,4,2,1,696]
s: 699
a: [3,597,2,7,1,589]
s: 597
r: 190
r: 487
r: 492
r: 581
r: 586
r: 592
r: 632
r: 684
d: [2,737,4,0,0,5,52,205,200,195,511,597,143]
e: 737
d: [2,738,4,0,0,5,52,205,195,200,511,597,143]
e: 738
f: 738,675,654,589,574,551,539,536,519,504,495,475,469,330,126,127,128,129,130
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t23_ordinal3
% SZS output start Proof for ordinal3__t23_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r1_ordinal1(k10_ordinal2(X0,X1),k10_ordinal2(X0,X2)) => r1_ordinal1(X1,X2))))) [input t23_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r1_ordinal1(k10_ordinal2(X0,X1),k10_ordinal2(X0,X2)) => r1_ordinal1(X1,X2))))) [negated conjecture 1]
3. ! [X0,X1] : (r2_hidden(X0,X1) => ~r2_hidden(X1,X0)) [input antisymmetry_r2_hidden]
6. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
15. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1))) [input connectedness_r1_ordinal1]
18. ! [X0] : ! [X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [input d8_xboole_0]
19. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k10_ordinal2(X0,X1))) [input dt_k10_ordinal2]
34. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
37. ! [X0] : (v1_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_xboole_0(X0,X1) => r2_hidden(X0,X1)))) [input t11_ordinal1]
39. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (k10_ordinal2(X0,X1) = k10_ordinal2(X0,X2) => X1 = X2)))) [input t21_ordinal3]
40. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r2_hidden(k10_ordinal2(X0,X1),k10_ordinal2(X0,X2)) => r2_hidden(X1,X2))))) [input t22_ordinal3]
50. ! [X0,X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [flattening 18]
58. ! [X0,X1] : ((X0 != X1 & r1_tarski(X0,X1)) => r2_xboole_0(X0,X1)) [unused predicate definition removal 50]
60. ? [X0] : (? [X1] : (? [X2] : ((~r1_ordinal1(X1,X2) & r1_ordinal1(k10_ordinal2(X0,X1),k10_ordinal2(X0,X2))) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
61. ? [X0] : (? [X1] : (? [X2] : (~r1_ordinal1(X1,X2) & r1_ordinal1(k10_ordinal2(X0,X1),k10_ordinal2(X0,X2)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 60]
63. ! [X0] : (! [X1] : ((r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1)) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [ennf transformation 37]
64. ! [X0] : (! [X1] : (r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [flattening 63]
70. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 6]
71. ! [X0] : (! [X1] : (! [X2] : ((X1 = X2 | k10_ordinal2(X0,X1) != k10_ordinal2(X0,X2)) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 39]
72. ! [X0] : (! [X1] : (! [X2] : (X1 = X2 | k10_ordinal2(X0,X1) != k10_ordinal2(X0,X2) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 71]
73. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X1,X2) | ~r2_hidden(k10_ordinal2(X0,X1),k10_ordinal2(X0,X2))) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 40]
74. ! [X0] : (! [X1] : (! [X2] : (r2_hidden(X1,X2) | ~r2_hidden(k10_ordinal2(X0,X1),k10_ordinal2(X0,X2)) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 73]
83. ! [X0,X1] : (r2_xboole_0(X0,X1) | (X0 = X1 | ~r1_tarski(X0,X1))) [ennf transformation 58]
84. ! [X0,X1] : (r2_xboole_0(X0,X1) | X0 = X1 | ~r1_tarski(X0,X1)) [flattening 83]
91. ! [X0,X1] : (~r2_hidden(X1,X0) | ~r2_hidden(X0,X1)) [ennf transformation 3]
96. ! [X0,X1] : (v3_ordinal1(k10_ordinal2(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 19]
97. ! [X0,X1] : (v3_ordinal1(k10_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 96]
98. ! [X0,X1] : ((r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 15]
99. ! [X0,X1] : (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 98]
100. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 34]
101. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 100]
102. ? [X0] : (? [X1] : (? [X2] : (~r1_ordinal1(X1,X2) & r1_ordinal1(k10_ordinal2(X0,X1),k10_ordinal2(X0,X2)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (? [X2] : (~r1_ordinal1(X1,X2) & r1_ordinal1(k10_ordinal2(sK0,X1),k10_ordinal2(sK0,X2)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
103. ? [X1] : (? [X2] : (~r1_ordinal1(X1,X2) & r1_ordinal1(k10_ordinal2(sK0,X1),k10_ordinal2(sK0,X2)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) => (? [X2] : (~r1_ordinal1(sK1,X2) & r1_ordinal1(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,X2)) & v3_ordinal1(X2)) & v3_ordinal1(sK1)) [choice axiom]
104. ? [X2] : (~r1_ordinal1(sK1,X2) & r1_ordinal1(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,X2)) & v3_ordinal1(X2)) => (~r1_ordinal1(sK1,sK2) & r1_ordinal1(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) & v3_ordinal1(sK2)) [choice axiom]
105. ((~r1_ordinal1(sK1,sK2) & r1_ordinal1(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) & v3_ordinal1(sK2)) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 61,104,103,102]
113. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 101]
126. v3_ordinal1(sK0) [cnf transformation 105] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
127. v3_ordinal1(sK1) [cnf transformation 105] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
128. v3_ordinal1(sK2) [cnf transformation 105] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
129. r1_ordinal1(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) [cnf transformation 105] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
130. ~r1_ordinal1(sK1,sK2) [cnf transformation 105] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
135. ~r2_xboole_0(X0,X1) | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v1_ordinal1(X0) [cnf transformation 64] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
141. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 70] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
143. k10_ordinal2(X0,X1) != k10_ordinal2(X0,X2) | X1 = X2 | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 72] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
144. ~r2_hidden(k10_ordinal2(X0,X1),k10_ordinal2(X0,X2)) | r2_hidden(X1,X2) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 74] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
153. ~r1_tarski(X0,X1) | X0 = X1 | r2_xboole_0(X0,X1) [cnf transformation 84] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
164. ~r2_hidden(X1,X0) | ~r2_hidden(X0,X1) [cnf transformation 91] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
168. v3_ordinal1(k10_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 97] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
169. r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 99] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
170. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 113] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
183. 1 <=> r1_ordinal1(sK1,sK2) [avatar definition]
185. ~r1_ordinal1(sK1,sK2) <- (~1) [avatar component clause 183] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
186. ~1 [avatar split clause 130,183]
188. 2 <=> r1_ordinal1(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) [avatar definition]
190. r1_ordinal1(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) <- (2) [avatar component clause 188] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
191. 2 [avatar split clause 129,188]
193. 3 <=> v3_ordinal1(sK2) [avatar definition]
195. v3_ordinal1(sK2) <- (3) [avatar component clause 193] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
196. 3 [avatar split clause 128,193]
198. 4 <=> v3_ordinal1(sK1) [avatar definition]
200. v3_ordinal1(sK1) <- (4) [avatar component clause 198] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
201. 4 [avatar split clause 127,198]
203. 5 <=> v3_ordinal1(sK0) [avatar definition]
205. v3_ordinal1(sK0) <- (5) [avatar component clause 203] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
206. 5 [avatar split clause 126,203]
330. v1_ordinal1(sK2) <- (3) [resolution 141,195] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
346. 31 <=> v1_ordinal1(sK2) [avatar definition]
348. v1_ordinal1(sK2) <- (31) [avatar component clause 346] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
349. 31 | ~3 [avatar split clause 330,193,346]
465. r1_ordinal1(sK2,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK2) <- (~1) [resolution 169,185] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
468. r1_ordinal1(sK2,sK1) | ~v3_ordinal1(sK2) <- (~1, 4) [subsumption resolution 465,200] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
469. r1_ordinal1(sK2,sK1) <- (~1, 3, 4) [subsumption resolution 468,195] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
471. 45 <=> r1_ordinal1(sK2,sK1) [avatar definition]
473. r1_ordinal1(sK2,sK1) <- (45) [avatar component clause 471] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
474. 45 | 1 | ~3 | ~4 [avatar split clause 469,198,193,183,471]
475. r1_tarski(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) | ~v3_ordinal1(k10_ordinal2(sK0,sK2)) | ~v3_ordinal1(k10_ordinal2(sK0,sK1)) <- (2) [resolution 170,190] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
478. r1_tarski(sK2,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK2) <- (45) [resolution 170,473] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
482. 46 <=> v3_ordinal1(k10_ordinal2(sK0,sK1)) [avatar definition]
483. v3_ordinal1(k10_ordinal2(sK0,sK1)) <- (46) [avatar component clause 482] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
484. ~v3_ordinal1(k10_ordinal2(sK0,sK1)) <- (~46) [avatar component clause 482] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
486. 47 <=> v3_ordinal1(k10_ordinal2(sK0,sK2)) [avatar definition]
487. v3_ordinal1(k10_ordinal2(sK0,sK2)) <- (47) [avatar component clause 486] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
488. ~v3_ordinal1(k10_ordinal2(sK0,sK2)) <- (~47) [avatar component clause 486] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
490. 48 <=> r1_tarski(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) [avatar definition]
492. r1_tarski(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) <- (48) [avatar component clause 490] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
493. ~46 | ~47 | 48 | ~2 [avatar split clause 475,188,490,486,482]
494. r1_tarski(sK2,sK1) | ~v3_ordinal1(sK2) <- (4, 45) [subsumption resolution 478,200] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
495. r1_tarski(sK2,sK1) <- (3, 4, 45) [subsumption resolution 494,195] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
497. 49 <=> r1_tarski(sK2,sK1) [avatar definition]
499. r1_tarski(sK2,sK1) <- (49) [avatar component clause 497] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
500. 49 | ~3 | ~4 | ~45 [avatar split clause 495,471,198,193,497]
504. sK1 = sK2 | r2_xboole_0(sK2,sK1) <- (49) [resolution 499,153] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
506. 50 <=> r2_xboole_0(sK2,sK1) [avatar definition]
508. r2_xboole_0(sK2,sK1) <- (50) [avatar component clause 506] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
510. 51 <=> sK1 = sK2 [avatar definition]
511. sK1 != sK2 <- (~51) [avatar component clause 510] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
513. 50 | 51 | ~49 [avatar split clause 504,497,510,506]
516. r2_hidden(sK2,sK1) | ~v3_ordinal1(sK1) | ~v1_ordinal1(sK2) <- (50) [resolution 508,135] {a:4,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
518. r2_hidden(sK2,sK1) | ~v1_ordinal1(sK2) <- (4, 50) [subsumption resolution 516,200] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
519. r2_hidden(sK2,sK1) <- (4, 31, 50) [subsumption resolution 518,348] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
521. 52 <=> r2_hidden(sK2,sK1) [avatar definition]
523. r2_hidden(sK2,sK1) <- (52) [avatar component clause 521] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
524. 52 | ~4 | ~31 | ~50 [avatar split clause 519,506,346,198,521]
530. sK1 != sK2 | r1_ordinal1(sK1,sK2) | ~r1_ordinal1(sK2,sK1) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
532. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~46) [resolution 484,168] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
535. ~v3_ordinal1(sK0) <- (4, ~46) [subsumption resolution 532,200] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
536. $false <- (4, 5, ~46) [subsumption resolution 535,205] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
537. ~4 | ~5 | 46 [avatar contradiction clause 536]
539. v1_ordinal1(k10_ordinal2(sK0,sK1)) <- (46) [resolution 483,141] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
546. 55 <=> v1_ordinal1(k10_ordinal2(sK0,sK1)) [avatar definition]
548. v1_ordinal1(k10_ordinal2(sK0,sK1)) <- (55) [avatar component clause 546] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
549. 55 | ~46 [avatar split clause 539,482,546]
551. ~r2_hidden(sK1,sK2) <- (52) [resolution 523,164] {a:5,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
555. 56 <=> r2_hidden(sK1,sK2) [avatar definition]
557. ~r2_hidden(sK1,sK2) <- (~56) [avatar component clause 555] {a:5,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
558. ~56 | ~52 [avatar split clause 551,521,555]
570. ~v3_ordinal1(sK2) | ~v3_ordinal1(sK0) <- (~47) [resolution 488,168] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
573. ~v3_ordinal1(sK0) <- (3, ~47) [subsumption resolution 570,195] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
574. $false <- (3, 5, ~47) [subsumption resolution 573,205] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
575. ~3 | ~5 | 47 [avatar contradiction clause 574]
589. k10_ordinal2(sK0,sK1) = k10_ordinal2(sK0,sK2) | r2_xboole_0(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) <- (48) [resolution 492,153] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
591. 61 <=> r2_xboole_0(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) [avatar definition]
593. r2_xboole_0(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) <- (61) [avatar component clause 591] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
595. 62 <=> k10_ordinal2(sK0,sK1) = k10_ordinal2(sK0,sK2) [avatar definition]
597. k10_ordinal2(sK0,sK1) = k10_ordinal2(sK0,sK2) <- (62) [avatar component clause 595] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
598. 61 | 62 | ~48 [avatar split clause 589,490,595,591]
651. r2_hidden(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) | ~v3_ordinal1(k10_ordinal2(sK0,sK2)) | ~v1_ordinal1(k10_ordinal2(sK0,sK1)) <- (61) [resolution 593,135] {a:3,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
653. r2_hidden(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) | ~v1_ordinal1(k10_ordinal2(sK0,sK1)) <- (47, 61) [subsumption resolution 651,487] {a:3,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
654. r2_hidden(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) <- (47, 55, 61) [subsumption resolution 653,548] {a:3,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
656. 67 <=> r2_hidden(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) [avatar definition]
658. r2_hidden(k10_ordinal2(sK0,sK1),k10_ordinal2(sK0,sK2)) <- (67) [avatar component clause 656] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
659. 67 | ~47 | ~55 | ~61 [avatar split clause 654,591,546,486,656]
666. r2_hidden(sK1,sK2) | ~v3_ordinal1(sK2) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (67) [resolution 658,144] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
672. ~v3_ordinal1(sK2) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~56, 67) [subsumption resolution 666,557] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
673. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (3, ~56, 67) [subsumption resolution 672,195] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
674. ~v3_ordinal1(sK0) <- (3, 4, ~56, 67) [subsumption resolution 673,200] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
675. $false <- (3, 4, 5, ~56, 67) [subsumption resolution 674,205] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
676. ~3 | ~4 | ~5 | 56 | ~67 [avatar contradiction clause 675]
738. $false <- (3, 4, 5, ~51, 62) [unit resulting resolution 205,195,200,511,597,143] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:26,thDist:-26}
751. ~3 | ~4 | ~5 | 51 | ~62 [avatar contradiction clause 738]
760. $false [avatar sat refutation 186,191,196,201,206,349,474,493,500,513,524,530,537,549,558,575,598,659,676,751] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:580,thDist:-580}
% SZS output end Proof for ordinal3__t23_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 46
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 56
% Generated clauses: 340
% Active clauses: 123
% Passive clauses: 153
% Final active clauses: 109
% Final passive clauses: 28
% Main loop iterations started: 125
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 25
% Bw demodulations: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Equational tautologies: 2
% Forward subsumptions: 59
% 
% >>> Generating Inferences
% Binary resolution: 148
% Unit resulting resolution: 5
% Factoring: 1
% Forward superposition: 3
% Backward superposition: 12
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 17
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 80
% SAT solver unit clauses: 18
% SAT solver binary clauses: 40
% 
% Memory used [KB]: 5373
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

