i: [1,227,0,2,1,1,0,0] l6_midsp_3
a: [3,344,0,2,1,227]
s: 344
i: [1,228,0,3,1,1,0,0] l6_midsp_3
a: [3,339,0,3,1,228]
s: 339
i: [1,226,0,3,1,1,0,0] l6_midsp_3
a: [3,349,0,3,1,226]
s: 349
i: [1,225,0,3,1,1,0,0] l6_midsp_3
a: [3,354,0,3,1,225]
s: 354
i: [1,230,0,4,1,1,0,0] l6_midsp_3
a: [3,329,0,4,1,230]
s: 329
i: [1,229,0,6,1,1,0,0] l6_midsp_3
a: [3,334,0,6,1,229]
s: 334
i: [1,232,0,2,1,0,0,255] fc1_xboole_0
a: [3,359,0,2,1,232]
s: 359
i: [1,237,0,2,1,0,0,255] fc6_ordinal1
a: [3,379,0,2,1,237]
s: 379
i: [1,236,0,2,1,0,0,255] fc6_ordinal1
a: [3,384,0,2,1,236]
s: 384
i: [1,283,0,2,1,0,0,255] rc3_ordinal1
a: [3,389,0,2,1,283]
s: 389
i: [1,282,0,2,1,0,0,255] rc3_ordinal1
a: [3,394,0,2,1,282]
s: 394
i: [1,281,0,2,1,0,0,255] rc3_ordinal1
a: [3,399,0,2,1,281]
s: 399
i: [1,306,0,2,1,0,0,1] rc2_xboole_0
a: [3,404,0,2,1,306]
s: 404
i: [1,307,0,2,1,0,0,2] rc4_ordinal1
a: [3,409,0,2,1,307]
s: 409
i: [1,308,0,2,1,0,0,255] rc1_ordinal1
a: [3,414,0,2,1,308]
s: 414
i: [1,309,0,2,1,0,0,1] rc1_xboole_0
a: [3,419,0,2,1,309]
s: 419
i: [1,313,0,2,1,0,0,255] rc2_ordinal1
a: [3,424,0,2,1,313]
s: 424
i: [1,312,0,2,1,0,0,255] rc2_ordinal1
a: [3,429,0,2,1,312]
s: 429
i: [1,311,0,2,1,0,0,255] rc2_ordinal1
a: [3,434,0,2,1,311]
s: 434
i: [1,310,0,2,1,0,0,255] rc2_ordinal1
a: [3,439,0,2,1,310]
s: 439
i: [1,315,0,2,1,0,0,2] rc5_ordinal1
a: [3,444,0,2,1,315]
s: 444
i: [1,314,0,2,1,0,0,2] rc5_ordinal1
a: [3,449,0,2,1,314]
s: 449
i: [1,316,0,2,1,0,0,255] rc3_xreal_0
a: [3,454,0,2,1,316]
s: 454
i: [1,317,0,2,1,0,0,255] rc4_xreal_0
a: [3,459,0,2,1,317]
s: 459
i: [1,320,0,2,1,0,0,255] rc2_nat_1
a: [3,464,0,2,1,320]
s: 464
i: [1,319,0,2,1,0,0,255] rc2_nat_1
a: [3,469,0,2,1,319]
s: 469
i: [1,231,0,10,1,1,0,0] l6_midsp_3
a: [3,324,0,10,1,231]
s: 324
i: [1,285,0,3,1,0,0,2] rc1_finseq_2
s: 285
i: [1,286,0,3,1,0,0,255] reflexivity_r1_tarski
s: 286
i: [1,233,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,364,0,3,1,233]
s: 364
r: 384
r: 379
d: [2,476,0,2,1,2,23,384,364]
a: [3,481,0,2,1,476]
s: 481
d: [2,477,0,2,1,2,23,379,364]
a: [3,486,0,2,1,477]
s: 486
i: [1,234,0,3,1,0,0,255] t1_numerals
a: [3,369,0,3,1,234]
s: 369
d: [2,488,0,3,1,2,22,369,364]
a: [3,492,0,3,1,488]
s: 492
i: [1,239,0,4,2,0,0,255] cc6_ordinal1
s: 239
i: [1,243,0,4,2,0,0,255] cc3_nat_1
s: 243
d: [2,494,1,2,1,1,40,239,409]
a: [3,499,1,2,1,494]
s: 499
i: [1,244,0,4,2,0,0,255] cc1_nat_1
s: 244
d: [2,495,1,2,1,1,40,239,444]
a: [3,504,1,2,1,495]
s: 504
i: [1,247,0,4,2,0,0,255] cc1_ordinal1
s: 247
i: [1,248,0,4,2,0,0,255] cc1_ordinal1
s: 248
d: [2,506,1,2,1,1,40,243,454]
a: [3,510,1,2,1,506]
s: 510
i: [1,250,0,4,2,0,0,255] cc9_ordinal1
s: 250
d: [2,512,1,2,1,1,40,247,486]
a: [3,521,1,2,1,512]
s: 521
i: [1,251,0,4,2,0,0,255] cc4_ordinal1
s: 251
i: [1,252,0,4,2,0,0,2] cc7_ordinal1
s: 252
d: [2,514,1,2,1,1,40,247,414]
a: [3,531,1,2,1,514]
s: 531
i: [1,253,0,4,2,0,0,255] cc3_ordinal1
s: 253
d: [2,517,1,2,1,1,40,247,464]
a: [3,541,1,2,1,517]
s: 541
i: [1,277,0,4,1,0,0,255] existence_m1_finseq_1
s: 277
i: [1,278,0,4,1,0,0,2] existence_m2_finseq_1
s: 278
d: [2,543,1,2,1,1,40,248,486]
a: [3,552,1,2,1,543]
s: 552
i: [1,279,0,4,1,0,0,2] existence_m1_finseq_2
s: 279
d: [2,545,1,2,1,1,40,248,414]
a: [3,562,1,2,1,545]
s: 562
i: [1,280,0,4,1,0,0,1] existence_m1_subset_1
s: 280
i: [1,284,0,4,1,0,0,2] rc1_finseq_2
s: 284
d: [2,548,1,2,1,1,40,248,464]
a: [3,572,1,2,1,548]
s: 572
i: [1,299,0,4,1,0,0,255] rc1_relset_1
s: 299
d: [2,574,1,2,1,1,40,250,359]
a: [3,580,1,2,1,574]
s: 580
i: [1,300,0,4,1,0,0,255] rc1_relset_1
s: 300
i: [1,235,0,4,1,0,0,255] dt_k5_numbers
a: [3,374,0,4,1,235]
s: 374
d: [2,575,1,2,1,1,40,250,419]
a: [3,585,1,2,1,575]
s: 585
i: [1,318,0,4,1,0,0,255] rc2_nat_1
a: [3,474,0,4,1,318]
s: 474
d: [2,576,1,2,1,1,40,250,459]
a: [3,590,1,2,1,576]
s: 590
i: [1,254,0,5,2,0,0,255] t6_boole
s: 254
i: [1,255,0,5,2,0,0,255] cc8_ordinal1
s: 255
d: [2,592,1,2,1,1,40,251,359]
a: [3,598,1,2,1,592]
s: 598
i: [1,256,0,5,2,0,0,255] cc2_nat_1
s: 256
d: [2,657,1,2,1,1,40,256,354]
a: [3,668,1,2,1,657]
s: 668
i: [1,262,0,5,2,0,0,255] dt_m1_finseq_1
s: 262
i: [1,263,0,5,2,0,0,255] dt_m1_finseq_1
s: 263
d: [2,658,1,2,1,1,40,256,349]
a: [3,673,1,2,1,658]
s: 673
i: [1,264,0,5,2,0,0,255] dt_m1_finseq_1
s: 264
d: [2,593,1,2,1,1,40,251,419]
a: [3,603,1,2,1,593]
s: 603
i: [1,265,0,5,2,0,0,255] dt_m2_finseq_1
s: 265
i: [1,266,0,5,2,0,0,255] dt_m2_finseq_1
s: 266
d: [2,594,1,2,1,1,40,251,459]
a: [3,608,1,2,1,594]
s: 608
i: [1,274,0,5,2,0,0,1] t7_boole
s: 274
d: [2,685,1,3,1,1,40,274,329]
a: [3,689,1,3,1,685]
s: 689
d: [2,654,0,5,2,1,22,255,364]
s: 654
i: [1,257,0,6,3,0,0,255] cc2_ordinal1
s: 257
d: [2,692,1,2,1,2,40,654,354]
a: [3,699,1,2,1,692]
s: 699
i: [1,268,0,6,2,0,0,1] t1_subset
s: 268
d: [2,693,1,2,1,2,40,654,349]
a: [3,704,1,2,1,693]
s: 704
i: [1,270,0,6,2,0,0,255] redefinition_m2_finseq_1
s: 270
i: [1,271,0,6,2,0,0,255] redefinition_m2_finseq_1
s: 271
d: [2,715,2,2,1,1,40,699,239]
a: [3,719,2,2,1,715]
s: 719
i: [1,290,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 290
d: [2,727,2,2,1,1,40,704,239]
a: [3,731,2,2,1,727]
s: 731
i: [1,240,0,7,2,0,0,255] redefinition_k2_finseq_1
s: 240
i: [1,241,0,7,2,0,0,1] dt_k2_finseq_1
s: 241
d: [2,735,3,2,1,1,40,719,248]
a: [3,740,3,2,1,735]
s: 740
i: [1,249,0,7,3,0,0,255] cc5_ordinal1
s: 249
d: [2,771,1,4,2,1,40,249,339]
a: [3,787,1,2,1,771]
s: 787
i: [1,272,0,7,2,0,0,255] t3_subset
s: 272
i: [1,273,0,7,3,0,0,1] t8_boole
s: 273
d: [2,736,3,2,1,1,40,719,247]
a: [3,745,3,2,1,736]
s: 745
i: [1,287,0,7,2,0,0,1] dt_k4_finseq_2
s: 287
d: [2,753,3,2,1,1,40,731,248]
a: [3,758,3,2,1,753]
s: 758
i: [1,238,0,8,3,0,0,255] cc10_ordinal1
s: 238
i: [1,269,0,8,3,0,0,1] t2_subset
s: 269
d: [2,754,3,2,1,1,40,731,247]
a: [3,763,3,2,1,754]
s: 763
i: [1,296,0,8,3,0,0,1] fc5_finseq_2
s: 296
d: [2,822,1,5,2,1,40,269,354]
d: [2,850,1,3,1,2,21,822,481]
a: [3,854,1,3,1,850]
s: 854
i: [1,297,0,8,2,0,0,3] cc1_relset_1
s: 297
i: [1,298,0,8,1,0,0,255] rc1_relset_1
s: 298
d: [2,823,1,5,2,1,40,269,349]
d: [2,856,1,3,1,2,21,823,481]
a: [3,860,1,3,1,856]
s: 860
i: [1,258,0,9,4,0,0,255] fc4_ordinal1
s: 258
d: [2,824,1,5,2,1,40,269,339]
d: [2,862,1,3,1,2,21,824,344]
a: [3,866,1,3,1,862]
s: 866
i: [1,267,0,9,2,0,0,255] dt_m2_finseq_1
s: 267
i: [1,276,0,9,3,0,0,1] t5_subset
s: 276
d: [2,873,2,3,1,1,40,854,290]
a: [3,879,2,3,1,873]
s: 879
i: [1,294,0,9,2,0,0,1] existence_m2_finseq_2
s: 294
d: [2,887,2,3,1,1,40,860,290]
a: [3,893,2,3,1,887]
s: 893
i: [1,301,0,9,3,0,0,2] dt_k2_funct_7
s: 301
i: [1,302,0,9,3,0,0,2] dt_k2_funct_7
s: 302
d: [2,897,2,3,1,1,40,866,290]
a: [3,903,2,3,1,897]
s: 903
i: [1,275,0,10,3,0,0,1] t4_subset
s: 275
d: [2,721,1,4,1,1,40,268,329]
a: [3,725,1,4,1,721]
s: 725
i: [1,288,0,10,3,0,0,3] cc4_relset_1
s: 288
i: [1,289,0,10,3,0,0,3] cc3_relset_1
s: 289
d: [2,747,1,4,1,1,40,290,329]
a: [3,751,1,4,1,747]
s: 751
i: [1,291,0,10,3,0,0,1] dt_m2_finseq_2
s: 291
d: [2,945,1,8,2,1,40,291,334]
a: [3,951,1,5,1,945]
s: 951
d: [2,957,2,0,0,2,52,699,951,287]
e: 957
d: [2,958,2,2,1,1,40,951,287]
d: [2,960,2,0,0,2,21,958,699]
e: 960
r: 951
i: [1,292,0,10,3,0,0,1] redefinition_m2_finseq_2
s: 292
i: [1,293,0,10,3,0,0,1] redefinition_m2_finseq_2
s: 293
a: [3,955,1,3,1,945]
s: 955
i: [1,259,0,11,4,0,0,255] redefinition_k4_finseq_1
s: 259
d: [2,975,2,2,1,1,40,955,266]
a: [3,985,2,2,1,975]
s: 985
i: [1,260,0,11,4,0,0,255] dt_k4_finseq_1
s: 260
i: [1,303,0,11,2,0,0,2] fc8_relset_1
s: 303
d: [2,976,2,2,1,1,40,955,265]
a: [3,990,2,2,1,976]
s: 990
i: [1,295,0,12,4,0,0,255] cc1_finseq_2
s: 295
d: [2,974,2,3,1,1,40,955,270]
a: [3,980,2,3,1,974]
s: 980
i: [1,246,0,15,5,0,0,255] t124_finseq_2
s: 246
i: [1,261,0,16,4,0,0,1] t31_funct_7
s: 261
d: [2,1018,3,2,1,1,40,980,262]
a: [3,1022,3,2,1,1018]
s: 1022
i: [1,304,0,27,6,0,0,1] redefinition_k1_midsp_3
s: 304
a: [3,950,1,5,1,945]
s: 950
i: [1,305,0,27,6,0,0,1] dt_k1_midsp_3
s: 305
d: [2,962,1,10,2,1,40,292,334]
a: [3,968,1,5,1,962]
s: 968
d: [2,1044,2,7,3,1,40,968,295]
d: [2,1047,2,5,2,2,21,1044,699]
d: [2,1048,2,3,1,3,21,1047,344]
a: [3,1052,2,3,1,1048]
s: 1052
d: [2,1030,1,24,6,1,44,324,304]
d: [2,1031,1,21,5,2,21,1030,339]
d: [2,1032,1,18,4,3,21,1031,349]
d: [2,1033,1,15,3,4,21,1032,354]
d: [2,1034,1,13,2,5,21,1033,344]
a: [3,1038,1,8,1,1034]
s: 1038
d: [2,1046,2,6,2,1,40,968,249]
a: [3,1066,2,4,1,1046]
s: 1066
d: [2,610,1,2,1,1,40,252,359]
a: [3,616,1,2,1,610]
s: 616
d: [2,611,1,2,1,1,40,252,419]
a: [3,621,1,2,1,611]
s: 621
d: [2,1072,2,11,4,1,44,1038,261]
d: [2,1073,2,8,3,1,19,1072]
d: [2,1074,2,6,2,2,21,1073,990]
d: [2,1075,2,4,1,3,21,1074,1022]
a: [3,1079,2,4,1,1075]
s: 1079
d: [2,612,1,2,1,1,40,252,459]
a: [3,626,1,2,1,612]
s: 626
d: [2,1086,3,10,4,1,44,1079,259]
d: [2,1087,3,8,3,2,21,1086,985]
d: [2,1088,3,6,2,3,21,1087,990]
d: [2,1089,3,4,1,4,21,1088,1022]
a: [3,1093,3,4,1,1089]
s: 1093
d: [2,628,1,2,1,1,40,253,359]
a: [3,634,1,2,1,628]
s: 634
d: [2,629,1,2,1,1,40,253,419]
a: [3,639,1,2,1,629]
s: 639
d: [2,1045,2,9,2,1,40,968,269]
a: [3,1057,2,5,1,1045]
s: 1057
d: [2,630,1,2,1,1,40,253,459]
a: [3,644,1,2,1,630]
s: 644
a: [3,1060,2,4,1,1045]
s: 1060
d: [2,656,1,2,1,1,40,256,492]
a: [3,663,1,2,1,656]
s: 663
d: [2,796,1,4,2,1,40,238,374]
a: [3,803,1,2,1,796]
s: 803
d: [2,1130,3,5,1,1,40,1057,290]
a: [3,1136,3,5,1,1130]
s: 1136
a: [3,806,1,2,1,796]
s: 806
d: [2,1097,4,14,5,1,44,1093,246]
a: [3,1100,4,5,2,1097]
s: 1100
d: [2,1152,5,0,0,3,52,344,955,1100]
e: 1152
d: [2,1153,5,2,1,2,40,1100,955]
d: [2,1155,5,0,0,3,21,1153,344]
e: 1155
r: 1100
d: [2,911,1,5,2,1,40,276,374]
a: [3,922,1,2,1,911]
s: 922
d: [2,659,1,3,1,0,40,256,280]
a: [3,678,1,3,1,659]
s: 678
d: [2,513,2,2,1,1,40,247,499]
a: [3,526,2,2,1,513]
s: 526
d: [2,680,1,3,1,0,40,262,277]
s: 680
d: [2,516,2,2,1,1,40,247,504]
a: [3,536,2,2,1,516]
s: 536
d: [2,681,1,3,1,0,40,263,277]
s: 681
d: [2,682,1,3,1,0,40,264,277]
s: 682
d: [2,544,2,2,1,1,40,248,499]
a: [3,557,2,2,1,544]
s: 557
d: [2,683,1,3,1,0,40,265,278]
s: 683
d: [2,547,2,2,1,1,40,248,504]
a: [3,567,2,2,1,547]
s: 567
d: [2,684,1,3,1,0,40,266,278]
s: 684
d: [2,694,1,3,1,1,40,654,280]
a: [3,709,1,3,1,694]
s: 709
d: [2,1105,2,2,1,1,40,634,248]
a: [3,1111,2,2,1,1105]
s: 1111
d: [2,767,1,8,3,0,40,249,241]
a: [3,777,1,3,1,767]
s: 777
d: [2,1106,2,2,1,1,40,634,247]
a: [3,1116,2,2,1,1106]
s: 1116
d: [2,819,1,7,2,1,40,269,374]
a: [3,834,1,3,1,819]
s: 834
r: 834
a: [3,833,1,3,1,819]
s: 833
d: [2,1118,2,2,1,1,40,639,248]
a: [3,1123,2,2,1,1118]
s: 1123
d: [2,820,1,10,3,0,40,269,241]
a: [3,842,1,3,1,820]
s: 842
d: [2,1193,2,0,0,2,52,777,842,253]
e: 1193
d: [2,1194,2,3,1,1,40,842,253]
d: [2,1199,2,0,0,2,21,1194,777]
e: 1199
r: 842
d: [2,1119,2,2,1,1,40,639,247]
a: [3,1128,2,2,1,1119]
s: 1128
d: [2,821,1,5,2,1,40,269,492]
d: [2,844,1,3,1,2,21,821,481]
a: [3,848,1,3,1,844]
s: 848
d: [2,646,1,4,1,0,40,299,253]
s: 646
d: [2,1139,2,2,1,1,40,644,248]
a: [3,1144,2,2,1,1139]
s: 1144
d: [2,647,1,4,1,0,40,299,252]
s: 647
d: [2,1140,2,2,1,1,40,644,247]
a: [3,1149,2,2,1,1140]
s: 1149
d: [2,648,1,4,1,0,40,299,251]
s: 648
d: [2,649,1,4,1,0,40,299,250]
s: 649
a: [3,1103,4,9,3,1097]
s: 1103
d: [2,1214,5,0,0,4,52,699,1052,329,1103]
e: 1214
d: [2,1215,5,5,2,2,40,1103,329]
d: [2,1219,5,3,1,3,21,1215,699]
d: [2,1220,5,0,0,4,21,1219,1052]
e: 1220
f: 1220,1155,1097,1089,1075,1048,1034,1018,976,975,974,962,960,945,692,233,225,226,227,228,229,230,231
% Refutation found. Thanks to Tanya!
% SZS status Theorem for midsp_3__l6_midsp_3
% SZS output start Proof for midsp_3__l6_midsp_3
1. ! [X0] : (m1_subset_1(X0,k5_numbers) => ! [X1] : (m1_subset_1(X1,k5_numbers) => ! [X2] : (~v1_xboole_0(X2) => ! [X3] : (m1_subset_1(X3,X2) => ! [X4] : (m2_finseq_2(X4,X2,k4_finseq_2(X0,X2)) => (r2_hidden(X1,k2_finseq_1(X0)) => k1_funct_1(k1_midsp_3(X2,X0,X4,X1,X3),X1) = X3)))))) [input l6_midsp_3]
2. ~! [X0] : (m1_subset_1(X0,k5_numbers) => ! [X1] : (m1_subset_1(X1,k5_numbers) => ! [X2] : (~v1_xboole_0(X2) => ! [X3] : (m1_subset_1(X3,X2) => ! [X4] : (m2_finseq_2(X4,X2,k4_finseq_2(X0,X2)) => (r2_hidden(X1,k2_finseq_1(X0)) => k1_funct_1(k1_midsp_3(X2,X0,X4,X1,X3),X1) = X3)))))) [negated conjecture 1]
5. ! [X0,X1] : ((v7_ordinal1(X1) & ~v1_xboole_0(X0)) => ! [X2] : (m1_subset_1(X2,k4_finseq_2(X1,X0)) => v3_card_1(X2,X1))) [input cc1_finseq_2]
26. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
38. ! [X0,X1] : (v7_ordinal1(X0) => m1_finseq_2(k4_finseq_2(X0,X1),X1)) [input dt_k4_finseq_2]
42. ! [X0] : ! [X1] : (m1_finseq_1(X1,X0) => (v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1))) [input dt_m1_finseq_1]
45. ! [X0] : ! [X1] : (m2_finseq_1(X1,X0) => (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X0))) & v1_finseq_1(X1) & v1_funct_1(X1))) [input dt_m2_finseq_1]
46. ! [X0,X1] : (m1_finseq_2(X1,X0) => ! [X2] : (m2_finseq_2(X2,X0,X1) => m2_finseq_1(X2,X0))) [input dt_m2_finseq_2]
70. ! [X0,X1,X2,X3,X4] : ((m1_subset_1(X4,X0) & m1_subset_1(X3,k5_numbers) & m1_subset_1(X2,k4_finseq_2(X1,X0)) & m1_subset_1(X1,k5_numbers) & ~v1_xboole_0(X0)) => k1_midsp_3(X0,X1,X2,X3,X4) = k2_funct_7(X2,X3,X4)) [input redefinition_k1_midsp_3]
72. ! [X0] : ((v1_finseq_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => k4_finseq_1(X0) = k9_xtuple_0(X0)) [input redefinition_k4_finseq_1]
73. k5_numbers = k4_ordinal1 [input redefinition_k5_numbers]
74. ! [X0] : ! [X1] : (m2_finseq_1(X1,X0) <=> m1_finseq_1(X1,X0)) [input redefinition_m2_finseq_1]
75. ! [X0,X1] : (m1_finseq_2(X1,X0) => ! [X2] : (m2_finseq_2(X2,X0,X1) <=> m1_subset_1(X2,X1))) [input redefinition_m2_finseq_2]
77. ! [X0] : (v7_ordinal1(X0) => ! [X1] : (~v1_xboole_0(X1) => ! [X2] : ((m2_finseq_1(X2,X1) & v3_card_1(X2,X0)) => k2_finseq_1(X0) = k4_finseq_1(X2)))) [input t124_finseq_2]
81. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ! [X2] : (r2_hidden(X2,k9_xtuple_0(X0)) => k1_funct_1(k2_funct_7(X0,X2,X1),X2) = X1)) [input t31_funct_7]
88. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1,X2] : (r2_hidden(X2,k9_xtuple_0(X0)) => k1_funct_1(k2_funct_7(X0,X2,X1),X2) = X1)) [flattening 81]
89. ! [X0,X1] : (m1_finseq_1(X1,X0) => (v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1))) [flattening 42]
90. ! [X0,X1] : (m2_finseq_1(X1,X0) => (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X0))) & v1_finseq_1(X1) & v1_funct_1(X1))) [flattening 45]
95. ! [X0,X1] : (m2_finseq_1(X1,X0) <=> m1_finseq_1(X1,X0)) [flattening 74]
123. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : ((k1_funct_1(k1_midsp_3(X2,X0,X4,X1,X3),X1) != X3 & r2_hidden(X1,k2_finseq_1(X0))) & m2_finseq_2(X4,X2,k4_finseq_2(X0,X2))) & m1_subset_1(X3,X2)) & ~v1_xboole_0(X2)) & m1_subset_1(X1,k5_numbers)) & m1_subset_1(X0,k5_numbers)) [ennf transformation 2]
124. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (k1_funct_1(k1_midsp_3(X2,X0,X4,X1,X3),X1) != X3 & r2_hidden(X1,k2_finseq_1(X0)) & m2_finseq_2(X4,X2,k4_finseq_2(X0,X2))) & m1_subset_1(X3,X2)) & ~v1_xboole_0(X2)) & m1_subset_1(X1,k5_numbers)) & m1_subset_1(X0,k5_numbers)) [flattening 123]
131. ! [X0] : (! [X1] : (! [X2] : (k2_finseq_1(X0) = k4_finseq_1(X2) | (~m2_finseq_1(X2,X1) | ~v3_card_1(X2,X0))) | v1_xboole_0(X1)) | ~v7_ordinal1(X0)) [ennf transformation 77]
132. ! [X0] : (! [X1] : (! [X2] : (k2_finseq_1(X0) = k4_finseq_1(X2) | ~m2_finseq_1(X2,X1) | ~v3_card_1(X2,X0)) | v1_xboole_0(X1)) | ~v7_ordinal1(X0)) [flattening 131]
140. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 26]
146. ! [X0] : (k4_finseq_1(X0) = k9_xtuple_0(X0) | (~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 72]
147. ! [X0] : (k4_finseq_1(X0) = k9_xtuple_0(X0) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 146]
150. ! [X0] : (! [X1,X2] : (k1_funct_1(k2_funct_7(X0,X2,X1),X2) = X1 | ~r2_hidden(X2,k9_xtuple_0(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 88]
151. ! [X0] : (! [X1,X2] : (k1_funct_1(k2_funct_7(X0,X2,X1),X2) = X1 | ~r2_hidden(X2,k9_xtuple_0(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 150]
152. ! [X0,X1] : ((v1_finseq_1(X1) & v1_funct_1(X1) & v1_relat_1(X1)) | ~m1_finseq_1(X1,X0)) [ennf transformation 89]
153. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X0))) & v1_finseq_1(X1) & v1_funct_1(X1)) | ~m2_finseq_1(X1,X0)) [ennf transformation 90]
163. ! [X0,X1] : (m1_finseq_2(k4_finseq_2(X0,X1),X1) | ~v7_ordinal1(X0)) [ennf transformation 38]
167. ! [X0,X1] : (! [X2] : (m2_finseq_1(X2,X0) | ~m2_finseq_2(X2,X0,X1)) | ~m1_finseq_2(X1,X0)) [ennf transformation 46]
168. ! [X0,X1] : (! [X2] : (m2_finseq_2(X2,X0,X1) <=> m1_subset_1(X2,X1)) | ~m1_finseq_2(X1,X0)) [ennf transformation 75]
170. ! [X0,X1] : (! [X2] : (v3_card_1(X2,X1) | ~m1_subset_1(X2,k4_finseq_2(X1,X0))) | (~v7_ordinal1(X1) | v1_xboole_0(X0))) [ennf transformation 5]
171. ! [X0,X1] : (! [X2] : (v3_card_1(X2,X1) | ~m1_subset_1(X2,k4_finseq_2(X1,X0))) | ~v7_ordinal1(X1) | v1_xboole_0(X0)) [flattening 170]
178. ! [X0,X1,X2,X3,X4] : (k1_midsp_3(X0,X1,X2,X3,X4) = k2_funct_7(X2,X3,X4) | (~m1_subset_1(X4,X0) | ~m1_subset_1(X3,k5_numbers) | ~m1_subset_1(X2,k4_finseq_2(X1,X0)) | ~m1_subset_1(X1,k5_numbers) | v1_xboole_0(X0))) [ennf transformation 70]
179. ! [X0,X1,X2,X3,X4] : (k1_midsp_3(X0,X1,X2,X3,X4) = k2_funct_7(X2,X3,X4) | ~m1_subset_1(X4,X0) | ~m1_subset_1(X3,k5_numbers) | ~m1_subset_1(X2,k4_finseq_2(X1,X0)) | ~m1_subset_1(X1,k5_numbers) | v1_xboole_0(X0)) [flattening 178]
182. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (k1_funct_1(k1_midsp_3(X2,X0,X4,X1,X3),X1) != X3 & r2_hidden(X1,k2_finseq_1(X0)) & m2_finseq_2(X4,X2,k4_finseq_2(X0,X2))) & m1_subset_1(X3,X2)) & ~v1_xboole_0(X2)) & m1_subset_1(X1,k5_numbers)) & m1_subset_1(X0,k5_numbers)) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : (k1_funct_1(k1_midsp_3(X2,sK0,X4,X1,X3),X1) != X3 & r2_hidden(X1,k2_finseq_1(sK0)) & m2_finseq_2(X4,X2,k4_finseq_2(sK0,X2))) & m1_subset_1(X3,X2)) & ~v1_xboole_0(X2)) & m1_subset_1(X1,k5_numbers)) & m1_subset_1(sK0,k5_numbers)) [choice axiom]
183. ? [X1] : (? [X2] : (? [X3] : (? [X4] : (k1_funct_1(k1_midsp_3(X2,sK0,X4,X1,X3),X1) != X3 & r2_hidden(X1,k2_finseq_1(sK0)) & m2_finseq_2(X4,X2,k4_finseq_2(sK0,X2))) & m1_subset_1(X3,X2)) & ~v1_xboole_0(X2)) & m1_subset_1(X1,k5_numbers)) => (? [X2] : (? [X3] : (? [X4] : (k1_funct_1(k1_midsp_3(X2,sK0,X4,sK1,X3),sK1) != X3 & r2_hidden(sK1,k2_finseq_1(sK0)) & m2_finseq_2(X4,X2,k4_finseq_2(sK0,X2))) & m1_subset_1(X3,X2)) & ~v1_xboole_0(X2)) & m1_subset_1(sK1,k5_numbers)) [choice axiom]
184. ? [X2] : (? [X3] : (? [X4] : (k1_funct_1(k1_midsp_3(X2,sK0,X4,sK1,X3),sK1) != X3 & r2_hidden(sK1,k2_finseq_1(sK0)) & m2_finseq_2(X4,X2,k4_finseq_2(sK0,X2))) & m1_subset_1(X3,X2)) & ~v1_xboole_0(X2)) => (? [X3] : (? [X4] : (k1_funct_1(k1_midsp_3(sK2,sK0,X4,sK1,X3),sK1) != X3 & r2_hidden(sK1,k2_finseq_1(sK0)) & m2_finseq_2(X4,sK2,k4_finseq_2(sK0,sK2))) & m1_subset_1(X3,sK2)) & ~v1_xboole_0(sK2)) [choice axiom]
185. ? [X3] : (? [X4] : (k1_funct_1(k1_midsp_3(sK2,sK0,X4,sK1,X3),sK1) != X3 & r2_hidden(sK1,k2_finseq_1(sK0)) & m2_finseq_2(X4,sK2,k4_finseq_2(sK0,sK2))) & m1_subset_1(X3,sK2)) => (? [X4] : (sK3 != k1_funct_1(k1_midsp_3(sK2,sK0,X4,sK1,sK3),sK1) & r2_hidden(sK1,k2_finseq_1(sK0)) & m2_finseq_2(X4,sK2,k4_finseq_2(sK0,sK2))) & m1_subset_1(sK3,sK2)) [choice axiom]
186. ? [X4] : (sK3 != k1_funct_1(k1_midsp_3(sK2,sK0,X4,sK1,sK3),sK1) & r2_hidden(sK1,k2_finseq_1(sK0)) & m2_finseq_2(X4,sK2,k4_finseq_2(sK0,sK2))) => (sK3 != k1_funct_1(k1_midsp_3(sK2,sK0,sK4,sK1,sK3),sK1) & r2_hidden(sK1,k2_finseq_1(sK0)) & m2_finseq_2(sK4,sK2,k4_finseq_2(sK0,sK2))) [choice axiom]
187. ((((sK3 != k1_funct_1(k1_midsp_3(sK2,sK0,sK4,sK1,sK3),sK1) & r2_hidden(sK1,k2_finseq_1(sK0)) & m2_finseq_2(sK4,sK2,k4_finseq_2(sK0,sK2))) & m1_subset_1(sK3,sK2)) & ~v1_xboole_0(sK2)) & m1_subset_1(sK1,k5_numbers)) & m1_subset_1(sK0,k5_numbers) [skolemisation 124,186,185,184,183,182]
188. ! [X0,X1] : ((m2_finseq_1(X1,X0) | ~m1_finseq_1(X1,X0)) & (m1_finseq_1(X1,X0) | ~m2_finseq_1(X1,X0))) [nnf transformation 95]
202. ! [X0,X1] : (! [X2] : ((m2_finseq_2(X2,X0,X1) | ~m1_subset_1(X2,X1)) & (m1_subset_1(X2,X1) | ~m2_finseq_2(X2,X0,X1))) | ~m1_finseq_2(X1,X0)) [nnf transformation 168]
225. m1_subset_1(sK0,k5_numbers) [cnf transformation 187] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
226. m1_subset_1(sK1,k5_numbers) [cnf transformation 187] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
227. ~v1_xboole_0(sK2) [cnf transformation 187] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
228. m1_subset_1(sK3,sK2) [cnf transformation 187] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
229. m2_finseq_2(sK4,sK2,k4_finseq_2(sK0,sK2)) [cnf transformation 187] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
230. r2_hidden(sK1,k2_finseq_1(sK0)) [cnf transformation 187] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
231. sK3 != k1_funct_1(k1_midsp_3(sK2,sK0,sK4,sK1,sK3),sK1) [cnf transformation 187] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
233. k5_numbers = k4_ordinal1 [cnf transformation 73] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
246. k2_finseq_1(X0) = k4_finseq_1(X2) | ~m2_finseq_1(X2,X1) | ~v3_card_1(X2,X0) | v1_xboole_0(X1) | ~v7_ordinal1(X0) [cnf transformation 132] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
255. v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1) [cnf transformation 140] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
259. k4_finseq_1(X0) = k9_xtuple_0(X0) | ~v1_finseq_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 147] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
261. k1_funct_1(k2_funct_7(X0,X2,X1),X2) = X1 | ~r2_hidden(X2,k9_xtuple_0(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 151] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
262. ~m1_finseq_1(X1,X0) | v1_relat_1(X1) [cnf transformation 152] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
265. ~m2_finseq_1(X1,X0) | v1_funct_1(X1) [cnf transformation 153] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
266. ~m2_finseq_1(X1,X0) | v1_finseq_1(X1) [cnf transformation 153] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
270. ~m2_finseq_1(X1,X0) | m1_finseq_1(X1,X0) [cnf transformation 188] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
287. m1_finseq_2(k4_finseq_2(X0,X1),X1) | ~v7_ordinal1(X0) [cnf transformation 163] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
291. ~m2_finseq_2(X2,X0,X1) | m2_finseq_1(X2,X0) | ~m1_finseq_2(X1,X0) [cnf transformation 167] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
292. ~m2_finseq_2(X2,X0,X1) | m1_subset_1(X2,X1) | ~m1_finseq_2(X1,X0) [cnf transformation 202] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
295. ~m1_subset_1(X2,k4_finseq_2(X1,X0)) | v3_card_1(X2,X1) | ~v7_ordinal1(X1) | v1_xboole_0(X0) [cnf transformation 171] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
304. k1_midsp_3(X0,X1,X2,X3,X4) = k2_funct_7(X2,X3,X4) | ~m1_subset_1(X4,X0) | ~m1_subset_1(X3,k5_numbers) | ~m1_subset_1(X2,k4_finseq_2(X1,X0)) | ~m1_subset_1(X1,k5_numbers) | v1_xboole_0(X0) [cnf transformation 179] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
322. 1 <=> sK3 = k1_funct_1(k1_midsp_3(sK2,sK0,sK4,sK1,sK3),sK1) [avatar definition]
324. sK3 != k1_funct_1(k1_midsp_3(sK2,sK0,sK4,sK1,sK3),sK1) <- (~1) [avatar component clause 322] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
325. ~1 [avatar split clause 231,322]
327. 2 <=> r2_hidden(sK1,k2_finseq_1(sK0)) [avatar definition]
329. r2_hidden(sK1,k2_finseq_1(sK0)) <- (2) [avatar component clause 327] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
330. 2 [avatar split clause 230,327]
332. 3 <=> m2_finseq_2(sK4,sK2,k4_finseq_2(sK0,sK2)) [avatar definition]
334. m2_finseq_2(sK4,sK2,k4_finseq_2(sK0,sK2)) <- (3) [avatar component clause 332] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
335. 3 [avatar split clause 229,332]
337. 4 <=> m1_subset_1(sK3,sK2) [avatar definition]
339. m1_subset_1(sK3,sK2) <- (4) [avatar component clause 337] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
340. 4 [avatar split clause 228,337]
342. 5 <=> v1_xboole_0(sK2) [avatar definition]
344. ~v1_xboole_0(sK2) <- (~5) [avatar component clause 342] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
345. ~5 [avatar split clause 227,342]
347. 6 <=> m1_subset_1(sK1,k5_numbers) [avatar definition]
349. m1_subset_1(sK1,k5_numbers) <- (6) [avatar component clause 347] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
350. 6 [avatar split clause 226,347]
352. 7 <=> m1_subset_1(sK0,k5_numbers) [avatar definition]
354. m1_subset_1(sK0,k5_numbers) <- (7) [avatar component clause 352] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
355. 7 [avatar split clause 225,352]
362. 9 <=> k5_numbers = k4_ordinal1 [avatar definition]
364. k5_numbers = k4_ordinal1 <- (9) [avatar component clause 362] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
365. 9 [avatar split clause 233,362]
654. ~m1_subset_1(X0,k5_numbers) | v7_ordinal1(X0) <- (9) [forward demodulation 255,364] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
692. v7_ordinal1(sK0) <- (7, 9) [resolution 654,354] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
697. 65 <=> v7_ordinal1(sK0) [avatar definition]
699. v7_ordinal1(sK0) <- (65) [avatar component clause 697] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
700. 65 | ~7 | ~9 [avatar split clause 692,362,352,697]
945. m2_finseq_1(sK4,sK2) | ~m1_finseq_2(k4_finseq_2(sK0,sK2),sK2) <- (3) [resolution 291,334] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
949. 99 <=> m1_finseq_2(k4_finseq_2(sK0,sK2),sK2) [avatar definition]
951. ~m1_finseq_2(k4_finseq_2(sK0,sK2),sK2) <- (~99) [avatar component clause 949] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
953. 100 <=> m2_finseq_1(sK4,sK2) [avatar definition]
955. m2_finseq_1(sK4,sK2) <- (100) [avatar component clause 953] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
956. ~99 | 100 | ~3 [avatar split clause 945,332,953,949]
958. ~v7_ordinal1(sK0) <- (~99) [resolution 951,287] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
960. $false <- (65, ~99) [subsumption resolution 958,699] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
961. ~65 | 99 [avatar contradiction clause 960]
962. m1_subset_1(sK4,k4_finseq_2(sK0,sK2)) | ~m1_finseq_2(k4_finseq_2(sK0,sK2),sK2) <- (3) [resolution 292,334] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
966. 101 <=> m1_subset_1(sK4,k4_finseq_2(sK0,sK2)) [avatar definition]
968. m1_subset_1(sK4,k4_finseq_2(sK0,sK2)) <- (101) [avatar component clause 966] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
969. ~99 | 101 | ~3 [avatar split clause 962,332,966,949]
974. m1_finseq_1(sK4,sK2) <- (100) [resolution 955,270] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
975. v1_finseq_1(sK4) <- (100) [resolution 955,266] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
976. v1_funct_1(sK4) <- (100) [resolution 955,265] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
978. 102 <=> m1_finseq_1(sK4,sK2) [avatar definition]
980. m1_finseq_1(sK4,sK2) <- (102) [avatar component clause 978] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
981. 102 | ~100 [avatar split clause 974,953,978]
983. 103 <=> v1_finseq_1(sK4) [avatar definition]
985. v1_finseq_1(sK4) <- (103) [avatar component clause 983] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
986. 103 | ~100 [avatar split clause 975,953,983]
988. 104 <=> v1_funct_1(sK4) [avatar definition]
990. v1_funct_1(sK4) <- (104) [avatar component clause 988] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
991. 104 | ~100 [avatar split clause 976,953,988]
1018. v1_relat_1(sK4) <- (102) [resolution 980,262] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1020. 108 <=> v1_relat_1(sK4) [avatar definition]
1022. v1_relat_1(sK4) <- (108) [avatar component clause 1020] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1023. 108 | ~102 [avatar split clause 1018,978,1020]
1030. sK3 != k1_funct_1(k2_funct_7(sK4,sK1,sK3),sK1) | ~m1_subset_1(sK3,sK2) | ~m1_subset_1(sK1,k5_numbers) | ~m1_subset_1(sK4,k4_finseq_2(sK0,sK2)) | ~m1_subset_1(sK0,k5_numbers) | v1_xboole_0(sK2) <- (~1) [superposition 324,304] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1031. sK3 != k1_funct_1(k2_funct_7(sK4,sK1,sK3),sK1) | ~m1_subset_1(sK1,k5_numbers) | ~m1_subset_1(sK4,k4_finseq_2(sK0,sK2)) | ~m1_subset_1(sK0,k5_numbers) | v1_xboole_0(sK2) <- (~1, 4) [subsumption resolution 1030,339] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1032. sK3 != k1_funct_1(k2_funct_7(sK4,sK1,sK3),sK1) | ~m1_subset_1(sK4,k4_finseq_2(sK0,sK2)) | ~m1_subset_1(sK0,k5_numbers) | v1_xboole_0(sK2) <- (~1, 4, 6) [subsumption resolution 1031,349] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1033. sK3 != k1_funct_1(k2_funct_7(sK4,sK1,sK3),sK1) | ~m1_subset_1(sK4,k4_finseq_2(sK0,sK2)) | v1_xboole_0(sK2) <- (~1, 4, 6, 7) [subsumption resolution 1032,354] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1034. sK3 != k1_funct_1(k2_funct_7(sK4,sK1,sK3),sK1) | ~m1_subset_1(sK4,k4_finseq_2(sK0,sK2)) <- (~1, 4, ~5, 6, 7) [subsumption resolution 1033,344] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1036. 109 <=> sK3 = k1_funct_1(k2_funct_7(sK4,sK1,sK3),sK1) [avatar definition]
1038. sK3 != k1_funct_1(k2_funct_7(sK4,sK1,sK3),sK1) <- (~109) [avatar component clause 1036] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1039. ~101 | ~109 | 1 | ~4 | 5 | ~6 | ~7 [avatar split clause 1034,352,347,342,337,322,1036,966]
1044. v3_card_1(sK4,sK0) | ~v7_ordinal1(sK0) | v1_xboole_0(sK2) <- (101) [resolution 968,295] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1047. v3_card_1(sK4,sK0) | v1_xboole_0(sK2) <- (65, 101) [subsumption resolution 1044,699] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1048. v3_card_1(sK4,sK0) <- (~5, 65, 101) [subsumption resolution 1047,344] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1050. 110 <=> v3_card_1(sK4,sK0) [avatar definition]
1052. v3_card_1(sK4,sK0) <- (110) [avatar component clause 1050] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1053. 110 | 5 | ~65 | ~101 [avatar split clause 1048,966,697,342,1050]
1072. sK3 != sK3 | ~r2_hidden(sK1,k9_xtuple_0(sK4)) | ~v1_funct_1(sK4) | ~v1_relat_1(sK4) <- (~109) [superposition 1038,261] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1073. ~r2_hidden(sK1,k9_xtuple_0(sK4)) | ~v1_funct_1(sK4) | ~v1_relat_1(sK4) <- (~109) [trivial inequality removal 1072] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1074. ~r2_hidden(sK1,k9_xtuple_0(sK4)) | ~v1_relat_1(sK4) <- (104, ~109) [subsumption resolution 1073,990] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1075. ~r2_hidden(sK1,k9_xtuple_0(sK4)) <- (104, 108, ~109) [subsumption resolution 1074,1022] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1077. 115 <=> r2_hidden(sK1,k9_xtuple_0(sK4)) [avatar definition]
1079. ~r2_hidden(sK1,k9_xtuple_0(sK4)) <- (~115) [avatar component clause 1077] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1080. ~115 | ~104 | ~108 | 109 [avatar split clause 1075,1036,1020,988,1077]
1086. ~r2_hidden(sK1,k4_finseq_1(sK4)) | ~v1_finseq_1(sK4) | ~v1_funct_1(sK4) | ~v1_relat_1(sK4) <- (~115) [superposition 1079,259] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1087. ~r2_hidden(sK1,k4_finseq_1(sK4)) | ~v1_funct_1(sK4) | ~v1_relat_1(sK4) <- (103, ~115) [subsumption resolution 1086,985] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1088. ~r2_hidden(sK1,k4_finseq_1(sK4)) | ~v1_relat_1(sK4) <- (103, 104, ~115) [subsumption resolution 1087,990] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1089. ~r2_hidden(sK1,k4_finseq_1(sK4)) <- (103, 104, 108, ~115) [subsumption resolution 1088,1022] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1091. 116 <=> r2_hidden(sK1,k4_finseq_1(sK4)) [avatar definition]
1093. ~r2_hidden(sK1,k4_finseq_1(sK4)) <- (~116) [avatar component clause 1091] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1094. ~116 | ~103 | ~104 | ~108 | 115 [avatar split clause 1089,1077,1020,988,983,1091]
1097. ~r2_hidden(sK1,k2_finseq_1(X0)) | ~m2_finseq_1(sK4,X1) | ~v3_card_1(sK4,X0) | v1_xboole_0(X1) | ~v7_ordinal1(X0) <- (~116) [superposition 1093,246] {a:4,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
1099. 117 <=> ! [X1] : (~m2_finseq_1(sK4,X1) | v1_xboole_0(X1)) [avatar definition]
1100. ~m2_finseq_1(sK4,X1) | v1_xboole_0(X1) <- (117) [avatar component clause 1099] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
1102. 118 <=> ! [X0] : (~r2_hidden(sK1,k2_finseq_1(X0)) | ~v7_ordinal1(X0) | ~v3_card_1(sK4,X0)) [avatar definition]
1103. ~r2_hidden(sK1,k2_finseq_1(X0)) | ~v7_ordinal1(X0) | ~v3_card_1(sK4,X0) <- (118) [avatar component clause 1102] {a:4,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
1104. 117 | 118 | 116 [avatar split clause 1097,1091,1102,1099]
1153. v1_xboole_0(sK2) <- (100, 117) [resolution 1100,955] {a:5,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:17,thDist:-17}
1155. $false <- (~5, 100, 117) [subsumption resolution 1153,344] {a:5,w:0,goal:1,sine:255,thAx:0,allAx:17,thDist:-17}
1156. 5 | ~100 | ~117 [avatar contradiction clause 1155]
1215. ~v7_ordinal1(sK0) | ~v3_card_1(sK4,sK0) <- (2, 118) [resolution 1103,329] {a:5,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:16,thDist:-16}
1219. ~v3_card_1(sK4,sK0) <- (2, 65, 118) [subsumption resolution 1215,699] {a:5,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:16,thDist:-16}
1220. $false <- (2, 65, 110, 118) [subsumption resolution 1219,1052] {a:5,w:0,goal:1,sine:255,thAx:0,allAx:16,thDist:-16}
1221. ~2 | ~65 | ~110 | ~118 [avatar contradiction clause 1220]
1222. $false [avatar sat refutation 325,330,335,340,345,350,355,365,700,956,961,969,981,986,991,1023,1039,1053,1080,1094,1104,1156,1221] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:904,thDist:-904}
% SZS output end Proof for midsp_3__l6_midsp_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 86
% >>> Preprocessing
% Introduced skolems: 22
% Pure predicates: 7
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 96
% Generated clauses: 468
% Active clauses: 184
% Passive clauses: 251
% Final active clauses: 178
% Final passive clauses: 64
% Main loop iterations started: 187
% 
% >>> Simplifying Inferences
% Duplicate literals: 14
% Trivial inequalities: 1
% Fw subsumption resolutions: 21
% Fw demodulations: 2
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 44
% 
% >>> Generating Inferences
% Binary resolution: 191
% Unit resulting resolution: 4
% Forward superposition: 12
% Backward superposition: 7
% Self superposition: 3
% 
% >>> AVATAR
% Split clauses: 19
% Split components: 38
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 137
% SAT solver unit clauses: 32
% SAT solver binary clauses: 72
% 
% Memory used [KB]: 10234
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% splitting model update: 0.001 s
% forward subsumption: 0.002 s
% other: 0.007 s ( own 0.003 s ) 

