i: [1,218,0,2,1,1,0,0] t123_xreal_1
a: [3,349,0,2,1,218]
s: 349
i: [1,222,0,7,2,1,0,0] t123_xreal_1
a: [3,343,0,3,1,222]
s: 343
a: [3,338,0,4,1,222]
s: 338
i: [1,223,0,2,1,0,0,2] fc1_numbers
a: [3,354,0,2,1,223]
s: 354
i: [1,224,0,2,1,0,0,2] fc1_xboole_0
a: [3,359,0,2,1,224]
s: 359
i: [1,231,0,2,1,0,0,3] fc6_ordinal1
a: [3,389,0,2,1,231]
s: 389
i: [1,230,0,2,1,0,0,3] fc6_ordinal1
a: [3,394,0,2,1,230]
s: 394
i: [1,301,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,398,0,2,1,301]
s: 398
i: [1,303,0,2,1,0,0,2] rc2_xboole_0
a: [3,406,0,2,1,303]
s: 406
i: [1,304,0,2,1,0,0,4] rc4_ordinal1
a: [3,411,0,2,1,304]
s: 411
i: [1,305,0,2,1,0,0,4] rc1_ordinal1
a: [3,416,0,2,1,305]
s: 416
i: [1,306,0,2,1,0,0,1] rc1_xreal_0
a: [3,421,0,2,1,306]
s: 421
i: [1,307,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,426,0,2,1,307]
s: 426
i: [1,308,0,2,1,0,0,2] rc1_xxreal_0
a: [3,431,0,2,1,308]
s: 431
d: [2,568,0,0,0,2,21,431,398]
e: 568
r: 398
s: 431
i: [1,309,0,2,1,0,0,2] rc1_xboole_0
a: [3,436,0,2,1,309]
s: 436
i: [1,313,0,2,1,0,0,255] rc2_ordinal1
a: [3,441,0,2,1,313]
s: 441
i: [1,312,0,2,1,0,0,255] rc2_ordinal1
a: [3,446,0,2,1,312]
s: 446
i: [1,311,0,2,1,0,0,255] rc2_ordinal1
a: [3,451,0,2,1,311]
s: 451
i: [1,310,0,2,1,0,0,255] rc2_ordinal1
a: [3,456,0,2,1,310]
s: 456
i: [1,315,0,2,1,0,0,4] rc5_ordinal1
a: [3,461,0,2,1,315]
s: 461
i: [1,314,0,2,1,0,0,4] rc5_ordinal1
a: [3,466,0,2,1,314]
s: 466
i: [1,317,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,471,0,2,1,317]
s: 471
i: [1,316,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,476,0,2,1,316]
s: 476
i: [1,321,0,2,1,0,0,2] rc2_xreal_0
a: [3,481,0,2,1,321]
s: 481
i: [1,320,0,2,1,0,0,2] rc2_xreal_0
a: [3,486,0,2,1,320]
s: 486
i: [1,319,0,2,1,0,0,2] rc2_xreal_0
a: [3,491,0,2,1,319]
s: 491
i: [1,318,0,2,1,0,0,2] rc2_xreal_0
a: [3,496,0,2,1,318]
s: 496
i: [1,325,0,2,1,0,0,2] rc3_xreal_0
a: [3,501,0,2,1,325]
s: 501
i: [1,324,0,2,1,0,0,2] rc3_xreal_0
a: [3,506,0,2,1,324]
s: 506
i: [1,323,0,2,1,0,0,2] rc3_xreal_0
a: [3,511,0,2,1,323]
s: 511
i: [1,322,0,2,1,0,0,2] rc3_xreal_0
a: [3,516,0,2,1,322]
s: 516
i: [1,327,0,2,1,0,0,2] rc3_xxreal_0
a: [3,521,0,2,1,327]
s: 521
i: [1,326,0,2,1,0,0,2] rc3_xxreal_0
a: [3,526,0,2,1,326]
s: 526
i: [1,329,0,2,1,0,0,2] rc2_xxreal_0
a: [3,531,0,2,1,329]
s: 531
i: [1,328,0,2,1,0,0,2] rc2_xxreal_0
a: [3,536,0,2,1,328]
s: 536
i: [1,333,0,2,1,0,0,2] rc4_xreal_0
a: [3,541,0,2,1,333]
s: 541
i: [1,332,0,2,1,0,0,2] rc4_xreal_0
a: [3,546,0,2,1,332]
s: 546
i: [1,331,0,2,1,0,0,2] rc4_xreal_0
a: [3,551,0,2,1,331]
s: 551
i: [1,330,0,2,1,0,0,2] rc4_xreal_0
a: [3,556,0,2,1,330]
s: 556
i: [1,335,0,2,1,0,0,2] rc4_xxreal_0
a: [3,561,0,2,1,335]
s: 561
i: [1,334,0,2,1,0,0,2] rc4_xxreal_0
a: [3,566,0,2,1,334]
s: 566
i: [1,295,0,3,1,0,0,255] reflexivity_r1_tarski
s: 295
i: [1,225,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,364,0,3,1,225]
s: 364
i: [1,226,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,369,0,3,1,226]
s: 369
r: 359
d: [2,570,0,2,1,2,23,359,369]
a: [3,574,0,2,1,570]
s: 574
i: [1,227,0,3,1,0,0,2] t1_numerals
a: [3,374,0,3,1,227]
s: 374
d: [2,576,0,3,1,2,22,374,369]
a: [3,580,0,3,1,576]
s: 580
i: [1,233,0,4,2,0,0,4] cc6_ordinal1
s: 233
i: [1,234,0,4,2,0,0,4] cc2_xreal_0
s: 234
i: [1,235,0,4,2,0,0,4] cc2_xcmplx_0
s: 235
d: [2,582,1,2,1,1,40,233,411]
a: [3,587,1,2,1,582]
s: 587
i: [1,236,0,4,2,0,0,4] cc2_xxreal_0
s: 236
d: [2,583,1,2,1,1,40,233,461]
a: [3,592,1,2,1,583]
s: 592
i: [1,237,0,4,2,0,0,255] cc1_ordinal1
s: 237
i: [1,238,0,4,2,0,0,255] cc1_ordinal1
s: 238
d: [2,594,1,2,1,1,40,234,411]
a: [3,599,1,2,1,594]
s: 599
i: [1,240,0,4,2,0,0,2] cc3_xreal_0
s: 240
d: [2,680,1,2,1,1,40,240,349]
a: [3,689,1,2,1,680]
s: 689
i: [1,241,0,4,2,0,0,2] cc4_xreal_0
s: 241
i: [1,254,0,4,2,0,0,255] cc9_ordinal1
s: 254
d: [2,697,1,2,1,1,40,241,349]
a: [3,706,1,2,1,697]
s: 706
i: [1,255,0,4,2,0,0,4] cc7_ordinal1
s: 255
d: [2,595,1,2,1,1,40,234,461]
a: [3,604,1,2,1,595]
s: 604
i: [1,256,0,4,2,0,0,4] cc3_ordinal1
s: 256
i: [1,294,0,4,1,0,0,3] existence_m1_subset_1
s: 294
d: [2,606,1,2,1,1,40,235,411]
a: [3,611,1,2,1,606]
s: 611
i: [1,228,0,4,1,0,0,2] dt_k5_numbers
a: [3,379,0,4,1,228]
s: 379
d: [2,607,1,2,1,1,40,235,461]
a: [3,616,1,2,1,607]
s: 616
i: [1,229,0,4,1,0,0,1] dt_k6_numbers
a: [3,384,0,4,1,229]
s: 384
d: [2,790,0,4,1,2,22,379,364]
a: [3,794,0,4,1,790]
s: 794
d: [2,618,1,2,1,1,40,236,411]
a: [3,623,1,2,1,618]
s: 623
d: [2,796,0,4,1,2,22,384,364]
a: [3,800,0,4,1,796]
s: 800
d: [2,619,1,2,1,1,40,236,461]
a: [3,628,1,2,1,619]
s: 628
i: [1,242,0,5,2,0,0,1] fc4_xreal_0
s: 242
i: [1,243,0,5,2,0,0,1] fc4_xreal_0
s: 243
d: [2,630,1,2,1,1,40,237,389]
a: [3,638,1,2,1,630]
s: 638
i: [1,252,0,5,2,0,0,1] dt_k5_xcmplx_0
s: 252
d: [2,632,1,2,1,1,40,237,416]
a: [3,648,1,2,1,632]
s: 648
i: [1,257,0,5,2,0,0,2] t6_boole
s: 257
i: [1,258,0,5,2,0,0,3] cc8_ordinal1
s: 258
d: [2,655,1,2,1,1,40,238,389]
a: [3,663,1,2,1,655]
s: 663
i: [1,259,0,5,2,0,0,2] cc1_xreal_0
s: 259
d: [2,657,1,2,1,1,40,238,416]
a: [3,673,1,2,1,657]
s: 673
i: [1,260,0,5,2,0,0,2] cc1_xcmplx_0
s: 260
i: [1,291,0,5,2,0,0,255] t7_boole
s: 291
d: [2,682,1,2,1,1,40,240,421]
a: [3,695,1,2,1,682]
s: 695
a: [3,401,0,5,2,301]
s: 401
d: [2,699,1,2,1,1,40,241,421]
a: [3,712,1,2,1,699]
s: 712
d: [2,804,0,5,2,1,22,257,369]
s: 804
i: [1,271,0,6,3,0,0,255] cc2_ordinal1
s: 271
d: [2,830,1,5,2,2,44,343,804]
s: 830
i: [1,278,0,6,3,0,0,2] cc3_xxreal_0
s: 278
d: [2,829,1,6,2,2,44,338,804]
s: 829
i: [1,280,0,6,3,0,0,2] cc3_xxreal_0
s: 280
i: [1,281,0,6,3,0,0,2] cc5_xxreal_0
s: 281
d: [2,848,2,6,2,3,40,829,401]
a: [3,852,2,3,1,848]
s: 852
i: [1,283,0,6,3,0,0,2] cc5_xxreal_0
s: 283
d: [2,714,1,2,1,1,40,254,574]
a: [3,721,1,2,1,714]
s: 721
i: [1,285,0,6,3,0,0,2] cc7_xxreal_0
s: 285
i: [1,286,0,6,3,0,0,2] cc7_xxreal_0
s: 286
d: [2,715,1,2,1,1,40,254,436]
a: [3,726,1,2,1,715]
s: 726
i: [1,287,0,6,2,0,0,255] t1_subset
s: 287
d: [2,716,1,2,1,1,40,254,556]
a: [3,731,1,2,1,716]
s: 731
i: [1,296,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 296
i: [1,239,0,7,3,0,0,4] cc5_ordinal1
s: 239
d: [2,717,1,2,1,1,40,254,566]
a: [3,736,1,2,1,717]
s: 736
i: [1,253,0,7,2,0,0,2] involutiveness_k5_xcmplx_0
s: 253
d: [2,738,1,2,1,1,40,255,574]
a: [3,745,1,2,1,738]
s: 745
i: [1,261,0,7,3,0,0,1] fc30_xreal_0
s: 261
i: [1,262,0,7,3,0,0,1] fc30_xreal_0
s: 262
d: [2,739,1,2,1,1,40,255,436]
a: [3,750,1,2,1,739]
s: 750
i: [1,263,0,7,3,0,0,1] fc28_xreal_0
s: 263
d: [2,740,1,2,1,1,40,255,556]
a: [3,755,1,2,1,740]
s: 755
i: [1,264,0,7,3,0,0,1] fc28_xreal_0
s: 264
i: [1,265,0,7,3,0,0,1] fc7_xcmplx_0
s: 265
d: [2,741,1,2,1,1,40,255,566]
a: [3,760,1,2,1,741]
s: 760
i: [1,266,0,7,3,0,0,1] fc7_xcmplx_0
s: 266
d: [2,766,1,2,1,1,40,256,574]
a: [3,773,1,2,1,766]
s: 773
i: [1,272,0,7,3,0,0,1] fc29_xreal_0
s: 272
i: [1,273,0,7,3,0,0,1] fc29_xreal_0
s: 273
d: [2,767,1,2,1,1,40,256,436]
a: [3,778,1,2,1,767]
s: 778
i: [1,274,0,7,3,0,0,1] fc27_xreal_0
s: 274
d: [2,768,1,2,1,1,40,256,556]
a: [3,783,1,2,1,768]
s: 783
i: [1,275,0,7,3,0,0,1] fc27_xreal_0
s: 275
i: [1,289,0,7,2,0,0,255] t3_subset
s: 289
d: [2,769,1,2,1,1,40,256,566]
a: [3,788,1,2,1,769]
s: 788
i: [1,290,0,7,3,0,0,2] t8_boole
s: 290
d: [2,834,1,4,2,1,40,278,486]
d: [2,836,1,2,1,2,21,834,491]
a: [3,840,1,2,1,836]
s: 840
i: [1,232,0,8,3,0,0,255] cc10_ordinal1
s: 232
i: [1,268,0,8,4,0,0,2] cc4_xxreal_0
s: 268
d: [2,1031,2,6,3,1,40,268,706]
a: [3,1045,2,2,1,1031]
s: 1045
r: 1045
i: [1,270,0,8,4,0,0,2] cc6_xxreal_0
s: 270
a: [3,1044,2,2,1,1031]
s: 1044
i: [1,276,0,8,4,0,0,2] cc8_xxreal_0
s: 276
i: [1,288,0,8,3,0,0,255] t2_subset
s: 288
a: [3,1049,2,2,1,1031]
s: 1049
i: [1,293,0,9,3,0,0,255] t5_subset
s: 293
d: [2,835,1,4,2,1,40,278,531]
d: [2,842,1,2,1,2,21,835,536]
a: [3,846,1,2,1,842]
s: 846
i: [1,292,0,10,3,0,0,255] t4_subset
s: 292
i: [1,302,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 302
d: [2,858,1,4,2,1,40,280,506]
d: [2,860,1,2,1,2,21,858,511]
a: [3,864,1,2,1,860]
s: 864
i: [1,246,0,11,5,0,0,1] t2_real
s: 246
a: [3,1052,2,2,1,1031]
s: 1052
i: [1,247,0,11,5,0,0,1] t1_real
s: 247
i: [1,248,0,11,5,0,0,1] t7_real
s: 248
d: [2,1177,1,10,4,1,40,248,338]
a: [3,1186,1,3,1,1177]
s: 1186
d: [2,1200,2,0,0,2,52,349,1186,243]
e: 1200
d: [2,1201,2,2,1,1,40,1186,243]
d: [2,1203,2,0,0,2,21,1201,349]
e: 1203
r: 1186
i: [1,249,0,11,5,0,0,1] t8_real
s: 249
a: [3,1185,1,3,1,1177]
s: 1185
d: [2,1205,3,0,0,2,52,852,1185,241]
e: 1205
d: [2,1206,2,3,1,1,40,1185,241]
d: [2,1209,2,0,0,2,21,1206,852]
e: 1209
r: 852
i: [1,250,0,11,5,0,0,1] t3_real
s: 250
i: [1,251,0,11,5,0,0,1] t4_real
s: 251
a: [3,1190,1,3,1,1177]
s: 1190
d: [2,1226,3,0,0,3,52,1052,349,1190,262]
e: 1226
d: [2,1227,2,4,2,1,40,1190,262]
d: [2,1231,2,2,1,2,21,1227,349]
d: [2,1232,2,0,0,3,21,1231,1052]
e: 1232
r: 1052
i: [1,244,0,13,6,0,0,1] t6_real
s: 244
d: [2,1152,1,8,4,1,40,246,343]
d: [2,1157,1,6,3,2,21,1152,349]
a: [3,1161,1,2,1,1157]
s: 1161
r: 1161
r: 1190
i: [1,245,0,13,6,0,0,1] t5_real
s: 245
i: [1,300,0,14,4,0,0,2] existence_m2_subset_1
s: 300
a: [3,1160,1,2,1,1157]
s: 1160
i: [1,297,0,15,5,0,0,2] dt_m2_subset_1
s: 297
a: [3,1194,1,2,1,1177]
s: 1194
r: 829
r: 338
r: 830
r: 343
r: 1194
a: [3,342,0,3,1,222]
s: 342
a: [3,339,0,4,1,222]
s: 339
a: [3,1193,1,2,1,1177]
s: 1193
i: [1,298,0,15,5,0,0,2] redefinition_m2_subset_1
s: 298
a: [3,1053,2,2,1,1031]
s: 1053
d: [2,1313,3,0,0,3,52,706,1049,1053,280]
e: 1313
d: [2,1315,3,4,2,1,40,1053,280]
d: [2,1317,3,2,1,2,21,1315,1049]
d: [2,1318,3,0,0,3,21,1317,706]
e: 1318
r: 1049
i: [1,299,0,15,5,0,0,2] redefinition_m2_subset_1
s: 299
a: [3,1189,1,3,1,1177]
s: 1189
d: [2,1327,3,0,0,3,52,1053,349,1189,273]
e: 1327
d: [2,1328,2,4,2,1,40,1189,273]
d: [2,1330,2,2,1,2,21,1328,349]
d: [2,1331,2,0,0,3,21,1330,1053]
e: 1331
i: [1,219,0,7,2,1,0,0] t123_xreal_1
d: [2,805,1,2,1,1,40,258,580]
d: [2,895,2,2,1,1,40,745,236]
d: [2,897,2,2,1,1,40,745,234]
d: [2,1255,2,4,2,1,40,1161,281]
d: [2,1258,2,2,1,2,21,1255,574]
d: [2,1281,2,4,2,1,40,1194,278]
d: [2,1283,2,2,1,2,21,1281,574]
d: [2,1285,1,8,4,1,40,342,248]
d: [2,1288,1,6,3,2,21,1285,349]
d: [2,1299,1,10,4,1,40,339,246]
d: [2,1306,1,8,3,2,21,1299,1160]
d: [2,1307,1,5,2,3,21,1306,1185]
f: 1331,1307,1288,1283,1258,1231,1203,1177,1157,897,895,805,576,570,227,226,224,218,219,222
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t123_xreal_1
% SZS output start Proof for xreal_1__t123_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0))) & ~(r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0)) & ~r1_xxreal_0(k6_numbers,X0)))) [input t123_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0))) & ~(r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0)) & ~r1_xxreal_0(k6_numbers,X0)))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
12. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
15. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
21. ! [X0] : ((v1_xreal_0(X0) & v3_xxreal_0(X0)) => (v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0)))) [input fc29_xreal_0]
22. ! [X0] : ((v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => (~v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0)))) [input fc30_xreal_0]
32. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
39. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v3_xxreal_0(X1) & ~v2_xxreal_0(X0) & ~r1_xxreal_0(X0,X1)))) [input t7_real]
51. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
60. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
73. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
81. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
86. ! [X0] : (v1_xreal_0(X0) => (v1_xreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0)))) [input fc4_xreal_0]
99. ? [X0] : (((r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0))) | (r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0)) & ~r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X0)) [ennf transformation 2]
102. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
104. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
109. ! [X0] : ((v1_xreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | ~v1_xreal_0(X0)) [ennf transformation 86]
114. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 32]
115. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 114]
118. ! [X0] : (! [X1] : ((v3_xxreal_0(X1) | v2_xxreal_0(X0) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 39]
119. ! [X0] : (! [X1] : (v3_xxreal_0(X1) | v2_xxreal_0(X0) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 118]
132. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 51]
135. ! [X0] : ((~v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | (~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 22]
136. ! [X0] : ((~v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 135]
147. ! [X0] : ((v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | (~v1_xreal_0(X0) | ~v3_xxreal_0(X0))) [ennf transformation 21]
148. ! [X0] : ((v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | ~v1_xreal_0(X0) | ~v3_xxreal_0(X0)) [flattening 147]
153. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 12]
154. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 153]
155. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 15]
156. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 155]
179. ? [X0] : (((r1_xxreal_0(k6_numbers,X0) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0))) | (r1_xxreal_0(k6_numbers,k5_xcmplx_0(X0)) & ~r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X0)) => (((r1_xxreal_0(k6_numbers,sK0) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0))) | (r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) & ~r1_xxreal_0(k6_numbers,sK0))) & v1_xreal_0(sK0)) [choice axiom]
180. ((r1_xxreal_0(k6_numbers,sK0) & ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0))) | (r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) & ~r1_xxreal_0(k6_numbers,sK0))) & v1_xreal_0(sK0) [skolemisation 99,179]
218. v1_xreal_0(sK0) [cnf transformation 180] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
219. ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) | ~r1_xxreal_0(k6_numbers,sK0) [cnf transformation 180] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
222. r1_xxreal_0(k6_numbers,sK0) | r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) [cnf transformation 180] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
224. v1_xboole_0(k1_xboole_0) [cnf transformation 73] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
226. k6_numbers = k1_xboole_0 [cnf transformation 81] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
227. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 60] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
234. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 102] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
236. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 104] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
243. v1_xreal_0(k5_xcmplx_0(X0)) | ~v1_xreal_0(X0) [cnf transformation 109] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
246. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 115] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
248. r1_xxreal_0(X0,X1) | v2_xxreal_0(X0) | v3_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 119] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
258. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 132] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
262. ~v3_xxreal_0(k5_xcmplx_0(X0)) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 136] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
273. v3_xxreal_0(k5_xcmplx_0(X0)) | ~v1_xreal_0(X0) | ~v3_xxreal_0(X0) [cnf transformation 148] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
278. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 154] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
281. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 156] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
337. 1 <=> r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) [avatar definition]
338. ~r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) <- (~1) [avatar component clause 337] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
339. r1_xxreal_0(k6_numbers,k5_xcmplx_0(sK0)) <- (1) [avatar component clause 337] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
341. 2 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
342. ~r1_xxreal_0(k6_numbers,sK0) <- (~2) [avatar component clause 341] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
343. r1_xxreal_0(k6_numbers,sK0) <- (2) [avatar component clause 341] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
344. 1 | 2 [avatar split clause 222,341,337]
345. ~2 | ~1 [avatar split clause 219,337,341]
347. 3 <=> v1_xreal_0(sK0) [avatar definition]
349. v1_xreal_0(sK0) <- (3) [avatar component clause 347] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
350. 3 [avatar split clause 218,347]
357. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
359. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 357] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
360. 5 [avatar split clause 224,357]
367. 7 <=> k6_numbers = k1_xboole_0 [avatar definition]
369. k6_numbers = k1_xboole_0 <- (7) [avatar component clause 367] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
370. 7 [avatar split clause 226,367]
372. 8 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
374. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (8) [avatar component clause 372] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
375. 8 [avatar split clause 227,372]
570. v1_xboole_0(k6_numbers) <- (5, 7) [backward demodulation 359,369] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
572. 48 <=> v1_xboole_0(k6_numbers) [avatar definition]
574. v1_xboole_0(k6_numbers) <- (48) [avatar component clause 572] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
575. 48 | ~5 | ~7 [avatar split clause 570,367,357,572]
576. m1_subset_1(k6_numbers,k4_ordinal1) <- (7, 8) [forward demodulation 374,369] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
578. 49 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
580. m1_subset_1(k6_numbers,k4_ordinal1) <- (49) [avatar component clause 578] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
581. 49 | ~7 | ~8 [avatar split clause 576,372,367,578]
743. 74 <=> v7_ordinal1(k6_numbers) [avatar definition]
745. v7_ordinal1(k6_numbers) <- (74) [avatar component clause 743] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
805. v7_ordinal1(k6_numbers) <- (49) [resolution 258,580] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
807. 74 | ~49 [avatar split clause 805,578,743]
895. v1_xxreal_0(k6_numbers) <- (74) [resolution 745,236] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
897. v1_xreal_0(k6_numbers) <- (74) [resolution 745,234] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
901. 95 <=> v1_xxreal_0(k6_numbers) [avatar definition]
904. 95 | ~74 [avatar split clause 895,743,901]
911. 97 <=> v1_xreal_0(k6_numbers) [avatar definition]
914. 97 | ~74 [avatar split clause 897,743,911]
1051. 115 <=> v3_xxreal_0(sK0) [avatar definition]
1053. v3_xxreal_0(sK0) <- (115) [avatar component clause 1051] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1152. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (2) [resolution 246,343] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1157. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 3) [subsumption resolution 1152,349] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1159. 133 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1160. ~v3_xxreal_0(k6_numbers) <- (~133) [avatar component clause 1159] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1161. v3_xxreal_0(k6_numbers) <- (133) [avatar component clause 1159] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1162. ~97 | 133 | ~115 | ~2 | ~3 [avatar split clause 1157,347,341,1051,1159,911]
1177. v2_xxreal_0(k6_numbers) | v3_xxreal_0(k5_xcmplx_0(sK0)) | ~v1_xreal_0(k5_xcmplx_0(sK0)) | ~v1_xreal_0(k6_numbers) <- (~1) [resolution 248,338] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1184. 135 <=> v1_xreal_0(k5_xcmplx_0(sK0)) [avatar definition]
1185. v1_xreal_0(k5_xcmplx_0(sK0)) <- (135) [avatar component clause 1184] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1186. ~v1_xreal_0(k5_xcmplx_0(sK0)) <- (~135) [avatar component clause 1184] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1188. 136 <=> v3_xxreal_0(k5_xcmplx_0(sK0)) [avatar definition]
1189. ~v3_xxreal_0(k5_xcmplx_0(sK0)) <- (~136) [avatar component clause 1188] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1190. v3_xxreal_0(k5_xcmplx_0(sK0)) <- (136) [avatar component clause 1188] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1192. 137 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1194. v2_xxreal_0(k6_numbers) <- (137) [avatar component clause 1192] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1195. ~97 | ~135 | 136 | 137 | 1 [avatar split clause 1177,337,1192,1188,1184,911]
1201. ~v1_xreal_0(sK0) <- (~135) [resolution 1186,243] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1203. $false <- (3, ~135) [subsumption resolution 1201,349] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1204. ~3 | 135 [avatar contradiction clause 1203]
1227. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (136) [resolution 1190,262] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1231. v3_xxreal_0(sK0) <- (3, 136) [subsumption resolution 1227,349] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1234. 115 | ~3 | ~136 [avatar split clause 1231,1188,347,1051]
1255. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (133) [resolution 1161,281] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1258. ~v1_xxreal_0(k6_numbers) <- (48, 133) [subsumption resolution 1255,574] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1259. ~95 | ~48 | ~133 [avatar split clause 1258,1159,572,901]
1281. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (137) [resolution 1194,278] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1283. ~v1_xxreal_0(k6_numbers) <- (48, 137) [subsumption resolution 1281,574] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1284. ~95 | ~48 | ~137 [avatar split clause 1283,1192,572,901]
1285. v2_xxreal_0(k6_numbers) | v3_xxreal_0(sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (~2) [resolution 342,248] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1288. v2_xxreal_0(k6_numbers) | v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (~2, 3) [subsumption resolution 1285,349] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1289. ~97 | 115 | 137 | 2 | ~3 [avatar split clause 1288,347,341,1192,1051,911]
1299. ~v3_xxreal_0(k5_xcmplx_0(sK0)) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k5_xcmplx_0(sK0)) | ~v1_xreal_0(k6_numbers) <- (1) [resolution 339,246] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1306. ~v3_xxreal_0(k5_xcmplx_0(sK0)) | ~v1_xreal_0(k5_xcmplx_0(sK0)) | ~v1_xreal_0(k6_numbers) <- (1, ~133) [subsumption resolution 1299,1160] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1307. ~v3_xxreal_0(k5_xcmplx_0(sK0)) | ~v1_xreal_0(k6_numbers) <- (1, ~133, 135) [subsumption resolution 1306,1185] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1308. ~97 | ~136 | ~1 | 133 | ~135 [avatar split clause 1307,1184,1159,337,1188,911]
1328. ~v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) <- (~136) [resolution 1189,273] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1330. ~v3_xxreal_0(sK0) <- (3, ~136) [subsumption resolution 1328,349] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1331. $false <- (3, 115, ~136) [subsumption resolution 1330,1053] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1332. ~3 | ~115 | 136 [avatar contradiction clause 1331]
1333. $false [avatar sat refutation 344,345,350,360,370,375,575,581,807,904,914,1162,1195,1204,1234,1259,1284,1289,1308,1332] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:820,thDist:-820}
% SZS output end Proof for xreal_1__t123_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 86
% >>> Preprocessing
% Introduced skolems: 19
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 118
% Generated clauses: 530
% Active clauses: 148
% Passive clauses: 251
% Final active clauses: 134
% Final passive clauses: 80
% Main loop iterations started: 166
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 40
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Forward subsumptions: 56
% 
% >>> Generating Inferences
% Binary resolution: 170
% Unit resulting resolution: 5
% Factoring: 1
% Forward superposition: 14
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 27
% Split components: 66
% Sat splitting refutations: 11
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 165
% SAT solver unit clauses: 46
% SAT solver binary clauses: 72
% 
% Memory used [KB]: 5628
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% congruence closure: 0.001 s
% forward subsumption: 0.002 s
% other: 0.005 s ( own 0.001 s ) 

