i: [1,68,0,3,1,1,0,0] t138_zfmisc_1
a: [3,103,0,3,1,68]
s: 103
i: [1,65,0,3,1,1,0,0] t138_zfmisc_1
a: [3,118,0,3,1,65]
s: 118
i: [1,67,0,6,1,1,0,0] t138_zfmisc_1
a: [3,108,0,6,1,67]
s: 108
i: [1,66,0,6,1,1,0,0] t138_zfmisc_1
a: [3,113,0,6,1,66]
s: 113
i: [1,69,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,123,0,2,1,69]
s: 123
i: [1,70,0,2,1,0,0,255] fc1_xboole_0
a: [3,128,0,2,1,70]
s: 128
i: [1,93,0,2,1,0,0,2] rc2_zfmisc_1
a: [3,138,0,2,1,93]
s: 138
i: [1,94,0,2,1,0,0,2] rc2_xboole_0
a: [3,143,0,2,1,94]
s: 143
i: [1,95,0,2,1,0,0,2] rc1_xboole_0
a: [3,148,0,2,1,95]
s: 148
i: [1,97,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,153,0,2,1,97]
s: 153
i: [1,96,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,158,0,2,1,96]
s: 158
i: [1,72,0,3,1,0,0,1] fc2_xboole_0
s: 72
i: [1,73,0,3,1,0,0,1] fc2_zfmisc_1
s: 73
i: [1,87,0,3,1,0,0,1] reflexivity_r1_tarski
s: 87
i: [1,98,0,3,1,0,0,1] d10_xboole_0
s: 98
i: [1,99,0,3,1,0,0,1] d10_xboole_0
s: 99
i: [1,71,0,3,1,0,0,255] d2_xboole_0
a: [3,133,0,3,1,71]
s: 133
r: 123
i: [1,75,0,4,2,0,0,2] cc2_zfmisc_1
s: 75
d: [2,161,1,2,1,1,40,75,148]
a: [3,166,1,2,1,161]
s: 166
i: [1,76,0,4,2,0,0,2] cc1_zfmisc_1
s: 76
i: [1,74,0,5,1,0,0,255] t1_boole
s: 74
d: [2,162,1,2,1,1,40,75,128]
a: [3,171,1,2,1,162]
s: 171
i: [1,77,0,5,2,0,0,255] t6_boole
s: 77
i: [1,84,0,5,2,0,0,255] t7_boole
s: 84
i: [1,88,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 88
i: [1,90,0,6,2,0,0,1] fc4_xboole_0
s: 90
i: [1,91,0,6,2,0,0,1] fc5_xboole_0
s: 91
i: [1,92,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 92
i: [1,81,0,7,2,0,0,255] l1_zfmisc_1
s: 81
i: [1,82,0,7,2,0,0,255] l1_zfmisc_1
s: 82
i: [1,83,0,7,3,0,0,1] t8_boole
s: 83
d: [2,181,1,4,1,0,40,81,87]
s: 181
i: [1,89,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 89
d: [2,183,2,4,1,0,40,181,92]
s: 183
i: [1,80,0,9,3,0,0,1] d10_xboole_0
s: 80
i: [1,85,0,11,3,0,0,1] t8_xboole_1
s: 85
d: [2,195,1,6,2,1,40,80,118]
d: [2,197,1,3,1,2,21,195,103]
a: [3,201,1,3,1,197]
s: 201
i: [1,86,0,12,3,0,0,255] t135_zfmisc_1
s: 86
d: [2,218,1,6,2,1,40,86,113]
d: [2,222,1,3,1,2,21,218,201]
a: [3,226,1,3,1,222]
s: 226
d: [2,196,1,12,2,1,40,80,113]
d: [2,203,1,6,1,2,21,196,108]
a: [3,207,1,6,1,203]
s: 207
d: [2,185,1,5,1,0,44,89,74]
s: 185
d: [2,229,2,2,1,1,40,226,84]
a: [3,238,2,2,1,229]
s: 238
d: [2,186,1,5,1,0,44,89,74]
s: 186
d: [2,228,2,3,1,1,40,226,92]
a: [3,233,2,3,1,228]
s: 233
d: [2,187,1,5,1,0,44,74,89]
s: 187
d: [2,190,1,5,1,0,44,74,89]
s: 190
d: [2,240,2,7,2,1,40,207,85]
d: [2,241,2,4,1,2,21,240,118]
a: [3,245,2,4,1,241]
s: 245
d: [2,256,3,0,0,2,52,226,245,82]
e: 256
d: [2,257,3,3,1,1,40,245,82]
d: [2,259,3,0,0,2,21,257,226]
e: 259
f: 259,241,222,203,197,65,66,67,68
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t138_zfmisc_1
% SZS output start Proof for zfmisc_1__t138_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ~(X0 != X1 & k2_xboole_0(X0,k1_tarski(X2)) != X1 & r1_tarski(X1,k2_xboole_0(X0,k1_tarski(X2))) & r1_tarski(X0,X1)) [input t138_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ~(X0 != X1 & k2_xboole_0(X0,k1_tarski(X2)) != X1 & r1_tarski(X1,k2_xboole_0(X0,k1_tarski(X2))) & r1_tarski(X0,X1)) [negated conjecture 1]
7. ! [X0] : ! [X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [input d10_xboole_0]
19. ! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [input l1_zfmisc_1]
25. ! [X0] : ! [X1] : ! [X2] : ~(~r1_tarski(X2,X0) & ~r2_hidden(X1,X2) & r1_tarski(X2,k2_xboole_0(X0,k1_tarski(X1)))) [input t135_zfmisc_1]
30. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [input t8_xboole_1]
31. ~! [X0,X1,X2] : ~(X0 != X1 & k2_xboole_0(X0,k1_tarski(X2)) != X1 & r1_tarski(X1,k2_xboole_0(X0,k1_tarski(X2))) & r1_tarski(X0,X1)) [flattening 2]
32. ! [X0,X1] : (X0 = X1 <=> (r1_tarski(X1,X0) & r1_tarski(X0,X1))) [flattening 7]
33. ! [X0,X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [flattening 19]
36. ! [X0,X1,X2] : ((r1_tarski(X2,X1) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),X1)) [flattening 30]
37. ! [X0,X1,X2] : ~(~r1_tarski(X2,X0) & ~r2_hidden(X1,X2) & r1_tarski(X2,k2_xboole_0(X0,k1_tarski(X1)))) [flattening 25]
40. ? [X0,X1,X2] : (X0 != X1 & k2_xboole_0(X0,k1_tarski(X2)) != X1 & r1_tarski(X1,k2_xboole_0(X0,k1_tarski(X2))) & r1_tarski(X0,X1)) [ennf transformation 31]
46. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | (~r1_tarski(X2,X1) | ~r1_tarski(X0,X1))) [ennf transformation 36]
47. ! [X0,X1,X2] : (r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1)) [flattening 46]
48. ! [X0,X1,X2] : (r1_tarski(X2,X0) | r2_hidden(X1,X2) | ~r1_tarski(X2,k2_xboole_0(X0,k1_tarski(X1)))) [ennf transformation 37]
52. ? [X0,X1,X2] : (X0 != X1 & k2_xboole_0(X0,k1_tarski(X2)) != X1 & r1_tarski(X1,k2_xboole_0(X0,k1_tarski(X2))) & r1_tarski(X0,X1)) => (sK0 != sK1 & sK1 != k2_xboole_0(sK0,k1_tarski(sK2)) & r1_tarski(sK1,k2_xboole_0(sK0,k1_tarski(sK2))) & r1_tarski(sK0,sK1)) [choice axiom]
53. sK0 != sK1 & sK1 != k2_xboole_0(sK0,k1_tarski(sK2)) & r1_tarski(sK1,k2_xboole_0(sK0,k1_tarski(sK2))) & r1_tarski(sK0,sK1) [skolemisation 40,52]
54. ! [X0,X1] : ((X0 = X1 | (~r1_tarski(X1,X0) | ~r1_tarski(X0,X1))) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [nnf transformation 32]
55. ! [X0,X1] : ((X0 = X1 | ~r1_tarski(X1,X0) | ~r1_tarski(X0,X1)) & ((r1_tarski(X1,X0) & r1_tarski(X0,X1)) | X0 != X1)) [flattening 54]
56. ! [X0,X1] : ((r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X1) | ~r1_tarski(k1_tarski(X0),X1))) [nnf transformation 33]
65. r1_tarski(sK0,sK1) [cnf transformation 53] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
66. r1_tarski(sK1,k2_xboole_0(sK0,k1_tarski(sK2))) [cnf transformation 53] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
67. sK1 != k2_xboole_0(sK0,k1_tarski(sK2)) [cnf transformation 53] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
68. sK0 != sK1 [cnf transformation 53] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
80. ~r1_tarski(X1,X0) | X0 = X1 | ~r1_tarski(X0,X1) [cnf transformation 55] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
82. r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 56] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
85. r1_tarski(k2_xboole_0(X0,X2),X1) | ~r1_tarski(X2,X1) | ~r1_tarski(X0,X1) [cnf transformation 47] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
86. ~r1_tarski(X2,k2_xboole_0(X0,k1_tarski(X1))) | r2_hidden(X1,X2) | r1_tarski(X2,X0) [cnf transformation 48] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
101. 1 <=> sK0 = sK1 [avatar definition]
103. sK0 != sK1 <- (~1) [avatar component clause 101] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
104. ~1 [avatar split clause 68,101]
106. 2 <=> sK1 = k2_xboole_0(sK0,k1_tarski(sK2)) [avatar definition]
108. sK1 != k2_xboole_0(sK0,k1_tarski(sK2)) <- (~2) [avatar component clause 106] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
109. ~2 [avatar split clause 67,106]
111. 3 <=> r1_tarski(sK1,k2_xboole_0(sK0,k1_tarski(sK2))) [avatar definition]
113. r1_tarski(sK1,k2_xboole_0(sK0,k1_tarski(sK2))) <- (3) [avatar component clause 111] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
114. 3 [avatar split clause 66,111]
116. 4 <=> r1_tarski(sK0,sK1) [avatar definition]
118. r1_tarski(sK0,sK1) <- (4) [avatar component clause 116] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
119. 4 [avatar split clause 65,116]
195. sK0 = sK1 | ~r1_tarski(sK1,sK0) <- (4) [resolution 80,118] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
196. sK1 = k2_xboole_0(sK0,k1_tarski(sK2)) | ~r1_tarski(k2_xboole_0(sK0,k1_tarski(sK2)),sK1) <- (3) [resolution 80,113] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
197. ~r1_tarski(sK1,sK0) <- (~1, 4) [subsumption resolution 195,103] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
199. 15 <=> r1_tarski(sK1,sK0) [avatar definition]
201. ~r1_tarski(sK1,sK0) <- (~15) [avatar component clause 199] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
202. ~15 | 1 | ~4 [avatar split clause 197,116,101,199]
203. ~r1_tarski(k2_xboole_0(sK0,k1_tarski(sK2)),sK1) <- (~2, 3) [subsumption resolution 196,108] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
205. 16 <=> r1_tarski(k2_xboole_0(sK0,k1_tarski(sK2)),sK1) [avatar definition]
207. ~r1_tarski(k2_xboole_0(sK0,k1_tarski(sK2)),sK1) <- (~16) [avatar component clause 205] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
208. ~16 | 2 | ~3 [avatar split clause 203,111,106,205]
218. r2_hidden(sK2,sK1) | r1_tarski(sK1,sK0) <- (3) [resolution 86,113] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
222. r2_hidden(sK2,sK1) <- (3, ~15) [subsumption resolution 218,201] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
224. 17 <=> r2_hidden(sK2,sK1) [avatar definition]
226. r2_hidden(sK2,sK1) <- (17) [avatar component clause 224] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
227. 17 | ~3 | 15 [avatar split clause 222,199,111,224]
240. ~r1_tarski(k1_tarski(sK2),sK1) | ~r1_tarski(sK0,sK1) <- (~16) [resolution 207,85] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
241. ~r1_tarski(k1_tarski(sK2),sK1) <- (4, ~16) [subsumption resolution 240,118] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
243. 20 <=> r1_tarski(k1_tarski(sK2),sK1) [avatar definition]
245. ~r1_tarski(k1_tarski(sK2),sK1) <- (~20) [avatar component clause 243] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
246. ~20 | ~4 | 16 [avatar split clause 241,205,116,243]
257. ~r2_hidden(sK2,sK1) <- (~20) [resolution 245,82] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
259. $false <- (17, ~20) [subsumption resolution 257,226] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
260. ~17 | 20 [avatar contradiction clause 259]
261. $false [avatar sat refutation 104,109,114,119,202,208,227,246,260] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:92,thDist:-92}
% SZS output end Proof for zfmisc_1__t138_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 33
% Generated clauses: 112
% Active clauses: 41
% Passive clauses: 58
% Final active clauses: 40
% Final passive clauses: 11
% Main loop iterations started: 47
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 1
% Forward subsumptions: 21
% 
% >>> Generating Inferences
% Binary resolution: 19
% Unit resulting resolution: 1
% Forward superposition: 16
% Backward superposition: 15
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 12
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

