i: [1,16,0,10,1,1,0,0] t31_xtuple_0
a: [3,24,0,10,1,16]
s: 24
i: [1,19,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 19
i: [1,17,0,6,1,0,0,1] d14_xtuple_0
s: 17
i: [1,20,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 20
i: [1,18,0,10,1,0,0,2] t23_xtuple_0
s: 18
d: [2,26,1,7,1,0,44,17,17]
s: 26
d: [2,37,2,8,1,0,44,17,26]
s: 37
d: [2,30,1,10,1,0,44,18,20]
s: 30
d: [2,52,2,9,1,0,44,30,18]
s: 52
d: [2,31,1,10,1,0,44,18,20]
s: 31
d: [2,33,1,10,1,0,44,20,18]
s: 33
d: [2,56,2,9,1,0,44,18,30]
s: 56
d: [2,34,1,10,1,0,44,20,18]
s: 34
d: [2,38,3,9,1,0,44,37,17]
s: 38
d: [2,27,1,11,1,0,44,18,17]
s: 27
d: [2,28,1,11,1,0,44,18,17]
s: 28
d: [2,45,3,9,1,0,44,17,37]
s: 45
d: [2,92,2,12,1,0,44,27,17]
d: [2,104,2,11,1,0,22,92,18]
d: [2,105,2,10,1,0,22,104,17]
s: 105
d: [2,138,3,9,1,1,44,24,105]
d: [2,144,3,0,0,1,19,138]
e: 144
d: [2,134,3,0,0,1,52,24,105]
e: 134
f: 134,16
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t31_xtuple_0
% SZS output start Proof for xtuple_0__t31_xtuple_0
1. ! [X0] : ! [X1] : k11_xtuple_0(k2_xboole_0(X0,X1)) = k2_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1)) [input t31_xtuple_0]
2. ~! [X0] : ! [X1] : k11_xtuple_0(k2_xboole_0(X0,X1)) = k2_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1)) [negated conjecture 1]
4. ! [X0] : k11_xtuple_0(X0) = k9_xtuple_0(k9_xtuple_0(X0)) [input d14_xtuple_0]
9. ! [X0] : ! [X1] : k9_xtuple_0(k2_xboole_0(X0,X1)) = k2_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) [input t23_xtuple_0]
10. ~! [X0,X1] : k11_xtuple_0(k2_xboole_0(X0,X1)) = k2_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1)) [flattening 2]
11. ! [X0,X1] : k9_xtuple_0(k2_xboole_0(X0,X1)) = k2_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) [flattening 9]
13. ? [X0,X1] : k11_xtuple_0(k2_xboole_0(X0,X1)) != k2_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1)) [ennf transformation 10]
14. ? [X0,X1] : k11_xtuple_0(k2_xboole_0(X0,X1)) != k2_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1)) => k11_xtuple_0(k2_xboole_0(sK0,sK1)) != k2_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) [choice axiom]
15. k11_xtuple_0(k2_xboole_0(sK0,sK1)) != k2_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) [skolemisation 13,14]
16. k11_xtuple_0(k2_xboole_0(sK0,sK1)) != k2_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) [cnf transformation 15] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
17. k11_xtuple_0(X0) = k9_xtuple_0(k9_xtuple_0(X0)) [cnf transformation 4] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
18. k9_xtuple_0(k2_xboole_0(X0,X1)) = k2_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) [cnf transformation 11] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
22. 1 <=> k11_xtuple_0(k2_xboole_0(sK0,sK1)) = k2_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) [avatar definition]
24. k11_xtuple_0(k2_xboole_0(sK0,sK1)) != k2_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1)) <- (~1) [avatar component clause 22] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
25. ~1 [avatar split clause 16,22]
27. k9_xtuple_0(k2_xboole_0(k9_xtuple_0(X0),X1)) = k2_xboole_0(k11_xtuple_0(X0),k9_xtuple_0(X1)) [superposition 18,17] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
92. k9_xtuple_0(k2_xboole_0(k9_xtuple_0(X1),k9_xtuple_0(X0))) = k2_xboole_0(k11_xtuple_0(X1),k11_xtuple_0(X0)) [superposition 27,17] {a:2,w:12,wCS:4800,sine:1,thAx:0,allAx:3,thDist:-3}
104. k2_xboole_0(k11_xtuple_0(X1),k11_xtuple_0(X0)) = k9_xtuple_0(k9_xtuple_0(k2_xboole_0(X1,X0))) [forward demodulation 92,18] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:3,thDist:-3}
105. k2_xboole_0(k11_xtuple_0(X1),k11_xtuple_0(X0)) = k11_xtuple_0(k2_xboole_0(X1,X0)) [forward demodulation 104,17] {a:2,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
134. $false <- (~1) [unit resulting resolution 24,105] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
146. 1 [avatar contradiction clause 134]
151. $false [avatar sat refutation 25,146] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:13,thDist:-13}
% SZS output end Proof for xtuple_0__t31_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 5
% Generated clauses: 129
% Active clauses: 13
% Passive clauses: 73
% Final active clauses: 13
% Final passive clauses: 55
% Main loop iterations started: 18
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw demodulations: 30
% 
% >>> Deletion Inferences
% Forward subsumptions: 11
% Fw demodulations to eq. taut.: 10
% 
% >>> Generating Inferences
% Unit resulting resolution: 1
% Forward superposition: 42
% Backward superposition: 47
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 5
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

