i: [1,134,0,15,3,1,0,0] t2_algseq_1
a: [3,195,0,4,1,134]
s: 195
i: [1,135,0,2,1,0,0,1] spc2_boole
a: [3,208,0,2,1,135]
s: 208
i: [1,136,0,2,1,0,0,1] spc1_boole
a: [3,213,0,2,1,136]
s: 213
i: [1,137,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,218,0,2,1,137]
s: 218
i: [1,147,0,2,1,0,0,2] fc6_ordinal1
a: [3,263,0,2,1,147]
s: 263
i: [1,146,0,2,1,0,0,2] fc6_ordinal1
a: [3,268,0,2,1,146]
s: 268
i: [1,183,0,2,1,0,0,2] rc4_ordinal1
a: [3,303,0,2,1,183]
s: 303
i: [1,184,0,2,1,0,0,3] rc1_ordinal1
a: [3,308,0,2,1,184]
s: 308
i: [1,188,0,2,1,0,0,255] rc2_ordinal1
a: [3,313,0,2,1,188]
s: 313
i: [1,187,0,2,1,0,0,255] rc2_ordinal1
a: [3,318,0,2,1,187]
s: 318
i: [1,186,0,2,1,0,0,255] rc2_ordinal1
a: [3,323,0,2,1,186]
s: 323
i: [1,185,0,2,1,0,0,255] rc2_ordinal1
a: [3,328,0,2,1,185]
s: 328
i: [1,190,0,2,1,0,0,2] rc5_ordinal1
a: [3,333,0,2,1,190]
s: 333
i: [1,189,0,2,1,0,0,2] rc5_ordinal1
a: [3,338,0,2,1,189]
s: 338
i: [1,191,0,2,1,0,0,255] rc4_xreal_0
a: [3,343,0,2,1,191]
s: 343
i: [1,176,0,3,1,0,0,255] reflexivity_r1_tarski
s: 176
i: [1,138,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,223,0,3,1,138]
s: 223
r: 195
d: [2,345,0,4,1,2,23,195,223]
a: [3,349,0,4,1,345]
s: 349
i: [1,139,0,3,1,0,0,255] d2_xboole_0
a: [3,228,0,3,1,139]
s: 228
i: [1,140,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,233,0,3,1,140]
s: 233
r: 268
r: 263
d: [2,357,0,2,1,2,23,268,233]
a: [3,362,0,2,1,357]
s: 362
d: [2,358,0,2,1,2,23,263,233]
a: [3,367,0,2,1,358]
s: 367
i: [1,141,0,3,1,0,0,2] t1_numerals
a: [3,238,0,3,1,141]
s: 238
i: [1,150,0,3,1,0,0,255] spc2_numerals
a: [3,273,0,3,1,150]
s: 273
i: [1,149,0,3,1,0,0,255] spc2_numerals
a: [3,278,0,3,1,149]
s: 278
i: [1,153,0,3,1,0,0,255] spc1_numerals
a: [3,288,0,3,1,153]
s: 288
i: [1,152,0,3,1,0,0,255] spc1_numerals
a: [3,293,0,3,1,152]
s: 293
d: [2,351,0,3,1,2,22,228,223]
a: [3,355,0,3,1,351]
s: 355
r: 218
d: [2,376,0,2,1,2,23,218,355]
a: [3,380,0,2,1,376]
s: 380
d: [2,369,0,3,1,2,22,238,233]
d: [2,370,0,3,1,3,22,369,223]
a: [3,374,0,3,1,370]
s: 374
i: [1,155,0,4,2,0,0,3] cc6_ordinal1
s: 155
d: [2,382,1,2,1,1,40,155,303]
a: [3,387,1,2,1,382]
s: 387
i: [1,159,0,4,2,0,0,255] cc1_ordinal1
s: 159
d: [2,383,1,2,1,1,40,155,333]
a: [3,392,1,2,1,383]
s: 392
i: [1,160,0,4,2,0,0,255] cc1_ordinal1
s: 160
i: [1,162,0,4,2,0,0,255] cc9_ordinal1
s: 162
d: [2,394,1,2,1,1,40,159,367]
a: [3,401,1,2,1,394]
s: 401
i: [1,163,0,4,2,0,0,2] cc7_ordinal1
s: 163
d: [2,396,1,2,1,1,40,159,308]
a: [3,411,1,2,1,396]
s: 411
i: [1,164,0,4,2,0,0,3] cc3_ordinal1
s: 164
i: [1,175,0,4,1,0,0,2] existence_m1_subset_1
s: 175
d: [2,419,1,2,1,1,40,160,367]
a: [3,427,1,2,1,419]
s: 427
i: [1,142,0,4,1,0,0,1] t49_card_1
a: [3,243,0,4,1,142]
s: 243
d: [2,421,1,2,1,1,40,160,308]
a: [3,437,1,2,1,421]
s: 437
i: [1,143,0,4,1,0,0,2] dt_k5_numbers
a: [3,248,0,4,1,143]
s: 248
i: [1,144,0,4,1,0,0,1] dt_k6_numbers
a: [3,253,0,4,1,144]
s: 253
d: [2,444,1,2,1,1,40,162,343]
a: [3,449,1,2,1,444]
s: 449
i: [1,148,0,4,1,0,0,255] spc2_numerals
a: [3,283,0,4,1,148]
s: 283
d: [2,445,1,2,1,1,40,162,380]
a: [3,454,1,2,1,445]
s: 454
i: [1,151,0,4,1,0,0,255] spc1_numerals
a: [3,298,0,4,1,151]
s: 298
d: [2,480,0,4,1,2,22,243,223]
a: [3,484,0,4,1,480]
s: 484
d: [2,456,1,2,1,1,40,163,343]
a: [3,461,1,2,1,456]
s: 461
i: [1,165,0,5,2,0,0,1] t6_boole
s: 165
d: [2,457,1,2,1,1,40,163,380]
a: [3,466,1,2,1,457]
s: 466
i: [1,166,0,5,2,0,0,2] cc8_ordinal1
s: 166
i: [1,172,0,5,2,0,0,255] t7_boole
s: 172
d: [2,468,1,2,1,1,40,164,343]
a: [3,473,1,2,1,468]
s: 473
i: [1,145,0,5,1,0,0,1] t50_card_1
a: [3,258,0,5,1,145]
s: 258
d: [2,469,1,2,1,1,40,164,380]
a: [3,478,1,2,1,469]
s: 478
d: [2,488,0,5,2,1,22,165,223]
s: 488
d: [2,491,0,5,2,1,22,166,233]
s: 491
d: [2,524,1,6,2,2,44,349,488]
s: 524
d: [2,504,0,5,1,2,22,258,223]
a: [3,508,0,5,1,504]
s: 508
d: [2,532,1,2,1,2,40,491,293]
a: [3,538,1,2,1,532]
s: 538
i: [1,156,0,6,2,0,0,255] d6_card_1
s: 156
i: [1,167,0,6,3,0,0,255] cc2_ordinal1
s: 167
d: [2,533,1,2,1,2,40,491,278]
a: [3,543,1,2,1,533]
s: 543
i: [1,168,0,6,2,0,0,255] t1_subset
s: 168
d: [2,395,2,2,1,1,40,159,387]
a: [3,406,2,2,1,395]
s: 406
i: [1,178,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 178
i: [1,157,0,7,2,0,0,255] redefinition_k7_card_1
s: 157
d: [2,413,2,2,1,1,40,392,159]
a: [3,417,2,2,1,413]
s: 417
i: [1,158,0,7,2,0,0,1] dt_k7_card_1
s: 158
d: [2,420,2,2,1,1,40,160,387]
a: [3,432,2,2,1,420]
s: 432
i: [1,161,0,7,3,0,0,3] cc5_ordinal1
s: 161
i: [1,170,0,7,2,0,0,255] t3_subset
s: 170
d: [2,423,2,2,1,1,40,160,392]
a: [3,442,2,2,1,423]
s: 442
i: [1,171,0,7,3,0,0,1] t8_boole
s: 171
d: [2,492,2,2,1,1,40,473,160]
a: [3,497,2,2,1,492]
s: 497
i: [1,177,0,7,1,0,0,1] commutativity_k2_tarski
s: 177
d: [2,570,0,7,2,1,22,158,233]
s: 570
d: [2,493,2,2,1,1,40,473,159]
a: [3,502,2,2,1,493]
s: 502
i: [1,154,0,8,3,0,0,255] cc10_ordinal1
s: 154
d: [2,602,1,7,3,1,40,154,570]
a: [3,608,1,2,1,602]
s: 608
i: [1,169,0,8,3,0,0,255] t2_subset
s: 169
i: [1,174,0,9,3,0,0,255] t5_subset
s: 174
d: [2,603,1,4,2,1,40,154,248]
a: [3,616,1,2,1,603]
s: 616
i: [1,173,0,10,3,0,0,255] t4_subset
s: 173
d: [2,621,1,5,2,1,40,169,288]
a: [3,647,1,2,1,621]
s: 647
i: [1,182,0,14,4,0,0,2] existence_m2_subset_1
s: 182
i: [1,179,0,15,5,0,0,2] dt_m2_subset_1
s: 179
r: 647
a: [3,648,1,2,1,621]
s: 648
d: [2,702,2,0,0,2,52,616,648,162]
e: 702
d: [2,705,2,2,1,1,40,648,162]
d: [2,707,2,0,0,2,21,705,616]
e: 707
r: 608
r: 616
r: 648
i: [1,180,0,15,5,0,0,2] redefinition_m2_subset_1
s: 180
a: [3,607,1,2,1,602]
s: 607
i: [1,181,0,15,5,0,0,2] redefinition_m2_subset_1
s: 181
d: [2,573,1,4,2,1,40,161,288]
a: [3,581,1,2,1,573]
s: 581
r: 581
a: [3,582,1,2,1,573]
s: 582
a: [3,615,1,2,1,603]
s: 615
s: 647
d: [2,534,1,3,1,1,40,491,175]
a: [3,548,1,3,1,534]
s: 548
d: [2,592,1,8,3,1,40,570,161]
a: [3,596,1,3,1,592]
s: 596
d: [2,510,2,2,1,1,40,478,160]
a: [3,515,2,2,1,510]
s: 515
d: [2,619,1,5,2,1,40,169,374]
d: [2,627,1,3,1,2,21,619,362]
a: [3,631,1,3,1,627]
s: 631
d: [2,511,2,2,1,1,40,478,159]
a: [3,520,2,2,1,511]
s: 520
a: [3,644,1,3,1,621]
s: 644
d: [2,622,1,5,2,1,40,169,293]
d: [2,650,1,3,1,2,21,622,362]
a: [3,654,1,3,1,650]
s: 654
d: [2,551,2,2,1,1,40,538,155]
a: [3,555,2,2,1,551]
s: 555
d: [2,623,1,5,2,1,40,169,273]
a: [3,659,1,3,1,623]
s: 659
d: [2,560,2,2,1,1,40,543,155]
a: [3,564,2,2,1,560]
s: 564
d: [2,624,1,5,2,1,40,169,278]
d: [2,661,1,3,1,2,21,624,362]
a: [3,665,1,3,1,661]
s: 665
d: [2,678,1,5,2,1,40,174,248]
a: [3,682,1,3,1,678]
s: 682
d: [2,821,2,0,0,2,40,682,665]
e: 821
d: [2,820,2,0,0,2,40,682,654]
e: 820
d: [2,819,2,0,0,2,40,682,631]
e: 819
d: [2,817,2,0,0,2,52,631,682]
e: 817
d: [2,818,2,0,0,3,52,374,362,682,169]
e: 818
r: 682
d: [2,704,2,2,1,1,40,648,163]
a: [3,713,2,2,1,704]
s: 713
d: [2,827,3,0,0,2,52,582,713,155]
e: 827
d: [2,828,3,2,1,1,40,713,155]
d: [2,830,3,0,0,2,21,828,582]
e: 830
r: 713
d: [2,688,1,11,4,1,40,179,253]
d: [2,693,1,7,3,2,21,688,248]
d: [2,694,1,5,2,3,21,693,362]
d: [2,695,1,3,1,4,21,694,647]
a: [3,699,1,3,1,695]
s: 699
d: [2,777,3,2,1,1,40,555,160]
a: [3,782,3,2,1,777]
s: 782
d: [2,625,1,7,2,1,40,169,248]
a: [3,670,1,4,1,625]
s: 670
a: [3,673,1,3,1,625]
s: 673
d: [2,778,3,2,1,1,40,555,159]
a: [3,787,3,2,1,778]
s: 787
d: [2,525,1,5,2,2,44,374,488]
s: 525
d: [2,797,3,2,1,1,40,564,160]
a: [3,802,3,2,1,797]
s: 802
d: [2,578,1,5,2,0,40,161,175]
s: 578
a: [3,611,1,5,2,602]
s: 611
d: [2,798,3,2,1,1,40,564,159]
a: [3,807,3,2,1,798]
s: 807
d: [2,523,1,6,2,2,44,253,488]
s: 523
d: [2,743,2,3,1,1,40,548,155]
a: [3,747,2,3,1,743]
s: 747
d: [2,530,1,6,2,2,44,484,488]
s: 530
d: [2,567,1,8,3,0,44,156,157]
d: [2,568,1,6,2,0,18,567]
s: 568
d: [2,875,2,0,0,2,52,466,349,568]
e: 875
d: [2,876,2,5,2,1,44,349,568]
d: [2,883,2,2,1,1,19,876]
d: [2,885,2,0,0,2,21,883,466]
e: 885
r: 524
r: 349
a: [3,348,0,4,1,345]
s: 348
a: [3,199,0,5,1,134]
d: [2,888,0,4,1,2,22,199,484]
a: [3,892,0,4,1,888]
s: 892
d: [2,905,2,0,0,2,52,538,892,568]
e: 905
d: [2,906,2,5,2,1,44,892,568]
d: [2,907,2,2,1,1,19,906]
d: [2,909,2,0,0,2,21,907,538]
e: 909
r: 892
a: [3,891,0,4,1,888]
s: 891
r: 607
r: 611
r: 615
a: [3,203,0,6,1,134]
d: [2,911,0,4,1,2,22,203,508]
a: [3,915,0,4,1,911]
s: 915
d: [2,933,2,0,0,2,52,543,915,568]
e: 933
d: [2,934,2,5,2,1,44,915,568]
d: [2,935,2,2,1,1,19,934]
d: [2,937,2,0,0,2,21,935,543]
e: 937
f: 937,911,909,888,885,533,532,504,480,457,376,351,345,152,149,145,142,140,139,138,137,134
% Refutation found. Thanks to Tanya!
% SZS status Theorem for algseq_1__t2_algseq_1
% SZS output start Proof for algseq_1__t2_algseq_1
1. k7_card_1(2) = k2_tarski(k6_numbers,1) & k7_card_1(1) = k1_tarski(k6_numbers) & k7_card_1(k6_numbers) = k1_xboole_0 [input t2_algseq_1]
2. ~(k7_card_1(2) = k2_tarski(k6_numbers,1) & k7_card_1(1) = k1_tarski(k6_numbers) & k7_card_1(k6_numbers) = k1_xboole_0) [negated conjecture 1]
21. k5_numbers = k4_ordinal1 [input redefinition_k5_numbers]
30. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
36. ! [X0] : (v1_xboole_0(X0) => v7_ordinal1(X0)) [input cc7_ordinal1]
37. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
47. ! [X0] : (v7_ordinal1(X0) => k6_card_1(X0) = X0) [input d6_card_1]
49. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
50. ! [X0] : (v7_ordinal1(X0) => k6_card_1(X0) = k7_card_1(X0)) [input redefinition_k7_card_1]
57. k1_xboole_0 = o_0_0_xboole_0 [input d2_xboole_0]
58. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
59. m1_subset_1(2,k1_numbers) & m1_subset_1(2,k5_numbers) & m2_subset_1(2,k1_numbers,k5_numbers) & v2_xxreal_0(2) [input spc2_numerals]
62. 1 = k1_tarski(k1_xboole_0) [input t49_card_1]
63. 2 = k2_tarski(k1_xboole_0,1) [input t50_card_1]
74. m1_subset_1(2,k1_numbers) & m1_subset_1(2,k5_numbers) & m2_subset_1(2,k1_numbers,k5_numbers) [pure predicate removal 59]
75. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) [pure predicate removal 58]
88. k7_card_1(2) != k2_tarski(k6_numbers,1) | k7_card_1(1) != k1_tarski(k6_numbers) | k7_card_1(k6_numbers) != k1_xboole_0 [ennf transformation 2]
91. ! [X0] : (k6_card_1(X0) = X0 | ~v7_ordinal1(X0)) [ennf transformation 47]
92. ! [X0] : (k6_card_1(X0) = k7_card_1(X0) | ~v7_ordinal1(X0)) [ennf transformation 50]
97. ! [X0] : (v7_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 36]
100. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 37]
134. k7_card_1(2) != k2_tarski(k6_numbers,1) | k7_card_1(1) != k1_tarski(k6_numbers) | k7_card_1(k6_numbers) != k1_xboole_0 [cnf transformation 88] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
137. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 30] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
138. k6_numbers = k1_xboole_0 [cnf transformation 49] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
139. k1_xboole_0 = o_0_0_xboole_0 [cnf transformation 57] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
140. k5_numbers = k4_ordinal1 [cnf transformation 21] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
142. 1 = k1_tarski(k1_xboole_0) [cnf transformation 62] {a:0,w:4,wCS:1600,sine:1,thAx:0,allAx:1,thDist:-1}
145. 2 = k2_tarski(k1_xboole_0,1) [cnf transformation 63] {a:0,w:5,wCS:2000,sine:1,thAx:0,allAx:1,thDist:-1}
149. m1_subset_1(2,k5_numbers) [cnf transformation 74] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
152. m1_subset_1(1,k5_numbers) [cnf transformation 75] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
156. k6_card_1(X0) = X0 | ~v7_ordinal1(X0) [cnf transformation 91] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
157. k6_card_1(X0) = k7_card_1(X0) | ~v7_ordinal1(X0) [cnf transformation 92] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
163. ~v1_xboole_0(X0) | v7_ordinal1(X0) [cnf transformation 97] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
166. v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1) [cnf transformation 100] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
193. 1 <=> k7_card_1(k6_numbers) = k1_xboole_0 [avatar definition]
195. k7_card_1(k6_numbers) != k1_xboole_0 <- (~1) [avatar component clause 193] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
197. 2 <=> k7_card_1(1) = k1_tarski(k6_numbers) [avatar definition]
199. k7_card_1(1) != k1_tarski(k6_numbers) <- (~2) [avatar component clause 197] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
201. 3 <=> k7_card_1(2) = k2_tarski(k6_numbers,1) [avatar definition]
203. k7_card_1(2) != k2_tarski(k6_numbers,1) <- (~3) [avatar component clause 201] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
204. ~1 | ~2 | ~3 [avatar split clause 134,201,197,193]
216. 6 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
218. v1_xboole_0(o_0_0_xboole_0) <- (6) [avatar component clause 216] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
219. 6 [avatar split clause 137,216]
221. 7 <=> k6_numbers = k1_xboole_0 [avatar definition]
223. k6_numbers = k1_xboole_0 <- (7) [avatar component clause 221] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
224. 7 [avatar split clause 138,221]
226. 8 <=> k1_xboole_0 = o_0_0_xboole_0 [avatar definition]
228. k1_xboole_0 = o_0_0_xboole_0 <- (8) [avatar component clause 226] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
229. 8 [avatar split clause 139,226]
231. 9 <=> k5_numbers = k4_ordinal1 [avatar definition]
233. k5_numbers = k4_ordinal1 <- (9) [avatar component clause 231] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
234. 9 [avatar split clause 140,231]
241. 11 <=> 1 = k1_tarski(k1_xboole_0) [avatar definition]
243. 1 = k1_tarski(k1_xboole_0) <- (11) [avatar component clause 241] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
244. 11 [avatar split clause 142,241]
256. 14 <=> 2 = k2_tarski(k1_xboole_0,1) [avatar definition]
258. 2 = k2_tarski(k1_xboole_0,1) <- (14) [avatar component clause 256] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
259. 14 [avatar split clause 145,256]
276. 18 <=> m1_subset_1(2,k5_numbers) [avatar definition]
278. m1_subset_1(2,k5_numbers) <- (18) [avatar component clause 276] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
279. 18 [avatar split clause 149,276]
291. 21 <=> m1_subset_1(1,k5_numbers) [avatar definition]
293. m1_subset_1(1,k5_numbers) <- (21) [avatar component clause 291] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
294. 21 [avatar split clause 152,291]
345. k6_numbers != k7_card_1(k6_numbers) <- (~1, 7) [backward demodulation 195,223] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
347. 32 <=> k6_numbers = k7_card_1(k6_numbers) [avatar definition]
349. k6_numbers != k7_card_1(k6_numbers) <- (~32) [avatar component clause 347] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
350. ~32 | 1 | ~7 [avatar split clause 345,221,193,347]
351. k6_numbers = o_0_0_xboole_0 <- (7, 8) [forward demodulation 228,223] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
353. 33 <=> k6_numbers = o_0_0_xboole_0 [avatar definition]
355. k6_numbers = o_0_0_xboole_0 <- (33) [avatar component clause 353] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
356. 33 | ~7 | ~8 [avatar split clause 351,226,221,353]
376. v1_xboole_0(k6_numbers) <- (6, 33) [backward demodulation 218,355] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
378. 37 <=> v1_xboole_0(k6_numbers) [avatar definition]
380. v1_xboole_0(k6_numbers) <- (37) [avatar component clause 378] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
381. 37 | ~6 | ~33 [avatar split clause 376,353,216,378]
457. v7_ordinal1(k6_numbers) <- (37) [resolution 163,380] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
464. 51 <=> v7_ordinal1(k6_numbers) [avatar definition]
466. v7_ordinal1(k6_numbers) <- (51) [avatar component clause 464] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
467. 51 | ~37 [avatar split clause 457,378,464]
480. 1 = k1_tarski(k6_numbers) <- (7, 11) [forward demodulation 243,223] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
482. 54 <=> 1 = k1_tarski(k6_numbers) [avatar definition]
484. 1 = k1_tarski(k6_numbers) <- (54) [avatar component clause 482] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
485. 54 | ~7 | ~11 [avatar split clause 480,241,221,482]
491. ~m1_subset_1(X0,k5_numbers) | v7_ordinal1(X0) <- (9) [forward demodulation 166,233] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
504. 2 = k2_tarski(k6_numbers,1) <- (7, 14) [forward demodulation 258,223] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
506. 57 <=> 2 = k2_tarski(k6_numbers,1) [avatar definition]
508. 2 = k2_tarski(k6_numbers,1) <- (57) [avatar component clause 506] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
509. 57 | ~7 | ~14 [avatar split clause 504,256,221,506]
532. v7_ordinal1(1) <- (9, 21) [resolution 491,293] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
533. v7_ordinal1(2) <- (9, 18) [resolution 491,278] {a:1,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
536. 60 <=> v7_ordinal1(1) [avatar definition]
538. v7_ordinal1(1) <- (60) [avatar component clause 536] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
539. 60 | ~9 | ~21 [avatar split clause 532,291,231,536]
541. 61 <=> v7_ordinal1(2) [avatar definition]
543. v7_ordinal1(2) <- (61) [avatar component clause 541] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
544. 61 | ~9 | ~18 [avatar split clause 533,276,231,541]
567. k7_card_1(X0) = X0 | ~v7_ordinal1(X0) | ~v7_ordinal1(X0) [superposition 156,157] {a:1,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
568. k7_card_1(X0) = X0 | ~v7_ordinal1(X0) [duplicate literal removal 567] {a:1,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
876. k6_numbers != k6_numbers | ~v7_ordinal1(k6_numbers) <- (~32) [superposition 349,568] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
883. ~v7_ordinal1(k6_numbers) <- (~32) [trivial inequality removal 876] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
885. $false <- (~32, 51) [subsumption resolution 883,466] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
886. 32 | ~51 [avatar contradiction clause 885]
888. 1 != k7_card_1(1) <- (~2, 54) [forward demodulation 199,484] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
890. 101 <=> 1 = k7_card_1(1) [avatar definition]
892. 1 != k7_card_1(1) <- (~101) [avatar component clause 890] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
893. ~101 | 2 | ~54 [avatar split clause 888,482,197,890]
906. 1 != 1 | ~v7_ordinal1(1) <- (~101) [superposition 892,568] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
907. ~v7_ordinal1(1) <- (~101) [trivial inequality removal 906] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
909. $false <- (60, ~101) [subsumption resolution 907,538] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
910. ~60 | 101 [avatar contradiction clause 909]
911. 2 != k7_card_1(2) <- (~3, 57) [forward demodulation 203,508] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
913. 103 <=> 2 = k7_card_1(2) [avatar definition]
915. 2 != k7_card_1(2) <- (~103) [avatar component clause 913] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
916. ~103 | 3 | ~57 [avatar split clause 911,506,201,913]
934. 2 != 2 | ~v7_ordinal1(2) <- (~103) [superposition 915,568] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
935. ~v7_ordinal1(2) <- (~103) [trivial inequality removal 934] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
937. $false <- (61, ~103) [subsumption resolution 935,543] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
938. ~61 | 103 [avatar contradiction clause 937]
939. $false [avatar sat refutation 204,219,224,229,234,244,259,279,294,350,356,381,467,485,509,539,544,886,893,910,916,938] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:542,thDist:-542}
% SZS output end Proof for algseq_1__t2_algseq_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 62
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 5
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 58
% Generated clauses: 386
% Active clauses: 123
% Passive clauses: 178
% Final active clauses: 106
% Final passive clauses: 44
% Main loop iterations started: 130
% 
% >>> Simplifying Inferences
% Duplicate literals: 10
% Trivial inequalities: 4
% Fw subsumption resolutions: 24
% Fw demodulations: 12
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 39
% 
% >>> Generating Inferences
% Binary resolution: 124
% Unit resulting resolution: 7
% Forward superposition: 9
% Backward superposition: 17
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 13
% Split components: 27
% Sat splitting refutations: 15
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 126
% SAT solver unit clauses: 28
% SAT solver binary clauses: 58
% 
% Memory used [KB]: 5500
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% minimizing solver time: 0.001 s
% congruence closure: 0.002 s ( own 0.001 s ) 
% forward subsumption resolution: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

