i: [1,55,0,3,1,1,0,0] t91_zfmisc_1
a: [3,79,0,3,1,55]
s: 79
i: [1,54,0,3,1,1,0,0] t91_zfmisc_1
a: [3,84,0,3,1,54]
s: 84
i: [1,53,0,3,1,1,0,0] t91_zfmisc_1
a: [3,89,0,3,1,53]
s: 89
i: [1,52,0,7,1,1,0,0] t91_zfmisc_1
a: [3,94,0,7,1,52]
s: 94
i: [1,56,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,99,0,2,1,56]
s: 99
i: [1,57,0,2,1,0,0,1] fc1_xboole_0
a: [3,104,0,2,1,57]
s: 104
i: [1,74,0,2,1,0,0,2] rc2_xboole_0
a: [3,114,0,2,1,74]
s: 114
i: [1,75,0,2,1,0,0,2] rc1_xboole_0
a: [3,119,0,2,1,75]
s: 119
i: [1,59,0,3,1,0,0,255] fc2_xboole_0
s: 59
i: [1,58,0,3,1,0,0,255] d2_xboole_0
a: [3,109,0,3,1,58]
s: 109
r: 99
i: [1,70,0,4,1,0,0,2] fc1_zfmisc_1
s: 70
i: [1,71,0,4,1,0,0,255] fc3_xboole_0
s: 71
i: [1,60,0,5,2,0,0,1] t6_boole
s: 60
i: [1,66,0,5,2,0,0,2] t7_boole
s: 66
i: [1,73,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 73
i: [1,61,0,7,2,0,0,1] t7_xboole_0
s: 61
i: [1,65,0,7,3,0,0,1] t8_boole
s: 65
i: [1,72,0,7,1,0,0,255] commutativity_k2_tarski
s: 72
d: [2,122,1,7,3,0,44,60,60]
s: 122
i: [1,62,0,10,1,0,0,255] d5_tarski
s: 62
i: [1,67,0,10,2,0,0,1] l46_zfmisc_1
s: 67
d: [2,136,1,10,2,1,44,67,94]
s: 136
i: [1,68,0,10,2,0,0,1] l46_zfmisc_1
s: 68
d: [2,137,1,10,2,1,44,68,94]
s: 137
i: [1,63,0,13,3,0,0,2] t1_tarski
s: 63
i: [1,64,0,13,3,0,0,2] t1_tarski
s: 64
d: [2,124,1,7,2,0,40,61,73]
s: 124
i: [1,69,0,13,3,0,0,1] l46_zfmisc_1
s: 69
d: [2,145,2,9,3,1,40,69,137]
a: [3,155,2,6,2,145]
s: 155
r: 155
d: [2,149,1,13,3,1,44,69,94]
s: 149
d: [2,128,1,10,1,0,44,62,72]
s: 128
a: [3,152,2,3,1,145]
s: 152
r: 149
d: [2,195,3,0,0,2,52,89,152,61]
e: 195
d: [2,196,3,3,1,1,40,152,61]
d: [2,199,3,0,0,2,21,196,89]
e: 199
r: 152
s: 149
d: [2,177,2,10,3,1,40,149,66]
a: [3,183,2,4,1,177]
s: 183
d: [2,129,1,10,1,0,44,62,72]
s: 129
d: [2,131,1,10,1,0,44,62,72]
s: 131
d: [2,164,3,7,2,1,40,155,61]
d: [2,166,3,4,1,2,21,164,84]
a: [3,170,3,4,1,166]
s: 170
d: [2,132,1,10,1,0,44,62,72]
s: 132
d: [2,219,4,2,1,1,40,170,66]
a: [3,228,4,2,1,219]
s: 228
d: [2,133,1,10,1,0,44,72,62]
s: 133
d: [2,134,1,10,1,0,44,72,62]
s: 134
d: [2,218,4,4,1,1,40,170,73]
a: [3,223,4,4,1,218]
s: 223
d: [2,139,1,10,3,0,40,63,66]
s: 139
s: 155
d: [2,148,1,10,3,0,40,69,66]
s: 148
d: [2,138,1,13,3,0,40,63,73]
s: 138
d: [2,146,2,9,3,1,40,69,136]
a: [3,159,2,6,2,146]
s: 159
d: [2,147,1,13,3,0,40,69,73]
s: 147
d: [2,239,3,7,2,1,40,159,61]
d: [2,247,3,4,1,2,21,239,89]
a: [3,251,3,4,1,247]
s: 251
d: [2,130,1,15,1,0,44,62,62]
s: 130
d: [2,176,2,13,3,1,40,149,73]
s: 176
d: [2,256,4,2,1,1,40,251,66]
a: [3,270,4,2,1,256]
s: 270
d: [2,205,2,13,3,1,40,149,73]
s: 205
d: [2,238,4,4,1,2,40,159,170]
a: [3,245,4,4,1,238]
s: 245
d: [2,253,2,13,3,1,44,147,94]
s: 253
d: [2,189,2,10,1,0,44,128,72]
s: 189
d: [2,254,4,4,1,2,40,251,155]
a: [3,260,4,4,1,254]
s: 260
d: [2,190,2,10,1,0,44,128,72]
s: 190
d: [2,255,4,4,1,1,40,251,73]
a: [3,265,4,4,1,255]
s: 265
d: [2,192,2,10,1,0,44,72,128]
s: 192
d: [2,193,2,10,1,0,44,72,128]
s: 193
d: [2,283,5,4,1,1,40,245,73]
a: [3,288,5,4,1,283]
s: 288
d: [2,210,2,10,1,0,44,131,72]
s: 210
d: [2,303,5,4,1,1,40,260,73]
a: [3,308,5,4,1,303]
s: 308
d: [2,211,2,10,1,0,44,131,72]
s: 211
d: [2,231,2,10,3,0,40,139,73]
s: 231
d: [2,236,3,10,3,1,40,155,139]
s: 236
d: [2,187,2,14,1,0,44,128,62]
s: 187
d: [2,241,3,10,3,1,40,159,139]
s: 241
d: [2,188,2,14,1,0,44,128,128]
s: 188
d: [2,214,2,14,1,0,44,131,131]
s: 214
d: [2,312,4,10,3,1,40,236,73]
s: 312
d: [2,216,2,14,1,0,44,62,131]
s: 216
d: [2,327,4,10,3,1,40,241,73]
s: 327
d: [2,191,2,15,1,0,44,62,128]
s: 191
d: [2,212,2,15,1,0,44,131,62]
s: 212
d: [2,235,3,13,3,1,40,155,63]
s: 235
d: [2,213,2,15,1,0,44,131,128]
s: 213
d: [2,391,4,8,2,1,42,235]
d: [2,392,4,5,1,2,21,391,79]
a: [3,396,4,5,1,392]
s: 396
d: [2,215,2,15,1,0,44,128,131]
s: 215
d: [2,272,2,15,1,0,44,130,72]
s: 272
d: [2,423,5,5,1,2,40,396,159]
a: [3,429,5,5,1,423]
s: 429
d: [2,436,6,0,0,3,52,79,396,429,64]
e: 436
d: [2,437,6,8,2,1,40,429,64]
d: [2,442,6,5,1,2,21,437,79]
d: [2,443,6,0,0,3,21,442,396]
e: 443
r: 327
r: 241
r: 159
r: 429
a: [3,162,2,3,1,146]
s: 162
r: 149
r: 155
r: 176
d: [2,449,5,0,0,2,40,162,245]
e: 449
d: [2,448,4,0,0,2,40,162,251]
e: 448
d: [2,445,4,0,0,2,52,251,162]
e: 445
d: [2,446,3,0,0,2,52,84,162,61]
e: 446
d: [2,447,3,0,0,3,52,104,84,162,139]
e: 447
d: [2,450,3,3,1,1,40,162,61]
d: [2,458,3,0,0,2,21,450,84]
e: 458
d: [2,203,2,9,3,1,40,149,68]
d: [2,204,2,9,3,1,40,149,67]
f: 458,443,423,392,204,203,199,52,53,54,55
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t91_zfmisc_1
% SZS output start Proof for zfmisc_1__t91_zfmisc_1
1. ! [X0] : ! [X1] : (k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X1,X0) => (X0 = X1 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t91_zfmisc_1]
2. ~! [X0] : ! [X1] : (k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X1,X0) => (X0 = X1 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [negated conjecture 1]
17. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [input l46_zfmisc_1]
20. ! [X0] : ! [X1] : (! [X2] : (r2_hidden(X2,X0) <=> r2_hidden(X2,X1)) => X0 = X1) [input t1_tarski]
23. ! [X0] : ~(! [X1] : ~r2_hidden(X1,X0) & k1_xboole_0 != X0) [input t7_xboole_0]
25. ~! [X0,X1] : (k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X1,X0) => (X0 = X1 | k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 2]
27. ! [X0,X1] : (! [X2] : (r2_hidden(X2,X0) <=> r2_hidden(X2,X1)) => X0 = X1) [flattening 20]
30. ! [X0,X1,X2,X3] : (r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) <=> (r2_hidden(X1,X3) & r2_hidden(X0,X2))) [flattening 17]
31. ? [X0,X1] : ((X0 != X1 & k1_xboole_0 != X1 & k1_xboole_0 != X0) & k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X1,X0)) [ennf transformation 25]
32. ? [X0,X1] : (X0 != X1 & k1_xboole_0 != X1 & k1_xboole_0 != X0 & k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X1,X0)) [flattening 31]
34. ! [X0] : (? [X1] : r2_hidden(X1,X0) | k1_xboole_0 = X0) [ennf transformation 23]
35. ! [X0,X1] : (X0 = X1 | ? [X2] : (r2_hidden(X2,X0) <~> r2_hidden(X2,X1))) [ennf transformation 27]
39. ? [X0,X1] : (X0 != X1 & k1_xboole_0 != X1 & k1_xboole_0 != X0 & k2_zfmisc_1(X0,X1) = k2_zfmisc_1(X1,X0)) => (sK0 != sK1 & k1_xboole_0 != sK1 & k1_xboole_0 != sK0 & k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK1,sK0)) [choice axiom]
40. sK0 != sK1 & k1_xboole_0 != sK1 & k1_xboole_0 != sK0 & k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK1,sK0) [skolemisation 32,39]
41. ! [X0] : (? [X1] : r2_hidden(X1,X0) => r2_hidden(sK2(X0),X0)) [choice axiom]
42. ! [X0] : (r2_hidden(sK2(X0),X0) | k1_xboole_0 = X0) [skolemisation 34,41]
43. ! [X0,X1] : (X0 = X1 | ? [X2] : ((~r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) & (r2_hidden(X2,X1) | r2_hidden(X2,X0)))) [nnf transformation 35]
44. ! [X1,X0] : (? [X2] : ((~r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) & (r2_hidden(X2,X1) | r2_hidden(X2,X0))) => ((~r2_hidden(sK3(X0,X1),X1) | ~r2_hidden(sK3(X0,X1),X0)) & (r2_hidden(sK3(X0,X1),X1) | r2_hidden(sK3(X0,X1),X0)))) [choice axiom]
45. ! [X0,X1] : (X0 = X1 | ((~r2_hidden(sK3(X0,X1),X1) | ~r2_hidden(sK3(X0,X1),X0)) & (r2_hidden(sK3(X0,X1),X1) | r2_hidden(sK3(X0,X1),X0)))) [skolemisation 43,44]
46. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | (~r2_hidden(X1,X3) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [nnf transformation 30]
47. ! [X0,X1,X2,X3] : ((r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X3) & r2_hidden(X0,X2)) | ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)))) [flattening 46]
52. k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK1,sK0) [cnf transformation 40] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. k1_xboole_0 != sK0 [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. k1_xboole_0 != sK1 [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. sK0 != sK1 [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
61. r2_hidden(sK2(X0),X0) | k1_xboole_0 = X0 [cnf transformation 42] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
63. r2_hidden(sK3(X0,X1),X1) | X0 = X1 | r2_hidden(sK3(X0,X1),X0) [cnf transformation 45] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
64. ~r2_hidden(sK3(X0,X1),X1) | X0 = X1 | ~r2_hidden(sK3(X0,X1),X0) [cnf transformation 45] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
67. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X0,X2) [cnf transformation 47] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. ~r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | r2_hidden(X1,X3) [cnf transformation 47] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(X2,X3)) | ~r2_hidden(X1,X3) | ~r2_hidden(X0,X2) [cnf transformation 47] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
77. 1 <=> sK0 = sK1 [avatar definition]
79. sK0 != sK1 <- (~1) [avatar component clause 77] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. ~1 [avatar split clause 55,77]
82. 2 <=> k1_xboole_0 = sK1 [avatar definition]
84. k1_xboole_0 != sK1 <- (~2) [avatar component clause 82] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
85. ~2 [avatar split clause 54,82]
87. 3 <=> k1_xboole_0 = sK0 [avatar definition]
89. k1_xboole_0 != sK0 <- (~3) [avatar component clause 87] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
90. ~3 [avatar split clause 53,87]
92. 4 <=> k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK1,sK0) [avatar definition]
94. k2_zfmisc_1(sK0,sK1) = k2_zfmisc_1(sK1,sK0) <- (4) [avatar component clause 92] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
95. 4 [avatar split clause 52,92]
149. r2_hidden(k4_tarski(X0,X1),k2_zfmisc_1(sK0,sK1)) | ~r2_hidden(X1,sK0) | ~r2_hidden(X0,sK1) <- (4) [superposition 69,94] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
151. 10 <=> ! [X9] : ~r2_hidden(X9,sK0) [avatar definition]
152. ~r2_hidden(X9,sK0) <- (10) [avatar component clause 151] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
154. 11 <=> ! [X8] : (~r2_hidden(X8,sK1) | r2_hidden(X8,sK0)) [avatar definition]
155. ~r2_hidden(X8,sK1) | r2_hidden(X8,sK0) <- (11) [avatar component clause 154] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
158. 12 <=> ! [X11] : (~r2_hidden(X11,sK0) | r2_hidden(X11,sK1)) [avatar definition]
159. ~r2_hidden(X11,sK0) | r2_hidden(X11,sK1) <- (12) [avatar component clause 158] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
161. 13 <=> ! [X10] : ~r2_hidden(X10,sK1) [avatar definition]
162. ~r2_hidden(X10,sK1) <- (13) [avatar component clause 161] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
196. k1_xboole_0 = sK0 <- (10) [resolution 152,61] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
199. $false <- (~3, 10) [subsumption resolution 196,89] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
200. 3 | ~10 [avatar contradiction clause 199]
203. ~r2_hidden(X4,sK0) | ~r2_hidden(X5,sK1) | r2_hidden(X4,sK1) <- (4) [resolution 149,68] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
204. ~r2_hidden(X6,sK0) | ~r2_hidden(X7,sK1) | r2_hidden(X7,sK0) <- (4) [resolution 149,67] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
207. 13 | 12 | ~4 [avatar split clause 203,92,158,161]
208. 11 | 10 | ~4 [avatar split clause 204,92,151,154]
235. r2_hidden(sK3(X0,sK1),sK0) | sK1 = X0 | r2_hidden(sK3(X0,sK1),X0) <- (11) [resolution 155,63] {a:3,w:13,wCS:1300,nSel:1,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
391. r2_hidden(sK3(sK0,sK1),sK0) | sK0 = sK1 <- (11) [factoring 235] {a:4,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
392. r2_hidden(sK3(sK0,sK1),sK0) <- (~1, 11) [subsumption resolution 391,79] {a:4,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
394. 25 <=> r2_hidden(sK3(sK0,sK1),sK0) [avatar definition]
396. r2_hidden(sK3(sK0,sK1),sK0) <- (25) [avatar component clause 394] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
397. 25 | 1 | ~11 [avatar split clause 392,154,77,394]
423. r2_hidden(sK3(sK0,sK1),sK1) <- (12, 25) [resolution 396,159] {a:5,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
427. 26 <=> r2_hidden(sK3(sK0,sK1),sK1) [avatar definition]
429. r2_hidden(sK3(sK0,sK1),sK1) <- (26) [avatar component clause 427] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
430. 26 | ~12 | ~25 [avatar split clause 423,394,158,427]
437. sK0 = sK1 | ~r2_hidden(sK3(sK0,sK1),sK0) <- (26) [resolution 429,64] {a:6,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
442. ~r2_hidden(sK3(sK0,sK1),sK0) <- (~1, 26) [subsumption resolution 437,79] {a:6,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
443. $false <- (~1, 25, 26) [subsumption resolution 442,396] {a:6,w:0,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
444. 1 | ~25 | ~26 [avatar contradiction clause 443]
450. k1_xboole_0 = sK1 <- (13) [resolution 162,61] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
458. $false <- (~2, 13) [subsumption resolution 450,84] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
459. 2 | ~13 [avatar contradiction clause 458]
460. $false [avatar sat refutation 80,85,90,95,200,207,208,397,430,444,459] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:323,thDist:-323}
% SZS output end Proof for zfmisc_1__t91_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 294
% Active clauses: 68
% Passive clauses: 180
% Final active clauses: 57
% Final passive clauses: 95
% Main loop iterations started: 83
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 7
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Equational tautologies: 2
% Forward subsumptions: 64
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 71
% Unit resulting resolution: 5
% Factoring: 2
% Forward superposition: 71
% Backward superposition: 70
% Self superposition: 9
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 18
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 42
% SAT solver unit clauses: 9
% SAT solver binary clauses: 14
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward demodulation: 0.001 s
% superposition: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

