i: [1,125,0,2,1,1,0,0] t7_algseq_1
a: [3,198,0,2,1,125]
s: 198
i: [1,124,0,2,1,1,0,0] t7_algseq_1
a: [3,203,0,2,1,124]
s: 203
i: [1,127,0,3,1,1,0,0] t7_algseq_1
a: [3,184,0,3,1,127]
s: 184
i: [1,128,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,208,0,2,1,128]
s: 208
i: [1,132,0,2,1,0,0,1] fc6_ordinal1
a: [3,223,0,2,1,132]
s: 223
i: [1,131,0,2,1,0,0,1] fc6_ordinal1
a: [3,228,0,2,1,131]
s: 228
i: [1,163,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,232,0,2,1,163]
s: 232
i: [1,169,0,2,1,0,0,3] rc1_xreal_0
a: [3,240,0,2,1,169]
s: 240
i: [1,170,0,2,1,0,0,1] rc4_ordinal1
a: [3,245,0,2,1,170]
s: 245
i: [1,171,0,2,1,0,0,2] rc1_ordinal1
a: [3,250,0,2,1,171]
s: 250
i: [1,175,0,2,1,0,0,255] rc2_ordinal1
a: [3,255,0,2,1,175]
s: 255
i: [1,174,0,2,1,0,0,255] rc2_ordinal1
a: [3,260,0,2,1,174]
s: 260
i: [1,173,0,2,1,0,0,255] rc2_ordinal1
a: [3,265,0,2,1,173]
s: 265
i: [1,172,0,2,1,0,0,255] rc2_ordinal1
a: [3,270,0,2,1,172]
s: 270
i: [1,177,0,2,1,0,0,1] rc5_ordinal1
a: [3,275,0,2,1,177]
s: 275
i: [1,176,0,2,1,0,0,1] rc5_ordinal1
a: [3,280,0,2,1,176]
s: 280
i: [1,180,0,2,1,0,0,255] rc4_xreal_0
a: [3,285,0,2,1,180]
s: 285
i: [1,179,0,2,1,0,0,255] rc4_xreal_0
a: [3,290,0,2,1,179]
s: 290
d: [2,297,0,0,0,2,21,290,232]
e: 297
r: 232
s: 290
i: [1,178,0,2,1,0,0,255] rc4_xreal_0
a: [3,295,0,2,1,178]
s: 295
i: [1,126,0,18,2,1,0,0] t7_algseq_1
a: [3,193,0,9,1,126]
s: 193
i: [1,159,0,3,1,0,0,255] reflexivity_r1_tarski
s: 159
i: [1,129,0,3,1,0,0,255] d2_xboole_0
a: [3,213,0,3,1,129]
s: 213
i: [1,130,0,3,1,0,0,255] t1_numerals
a: [3,218,0,3,1,130]
s: 218
d: [2,299,0,3,1,2,22,218,213]
a: [3,303,0,3,1,299]
s: 303
i: [1,135,0,4,2,0,0,2] cc4_xreal_0
s: 135
d: [2,305,1,2,1,1,40,135,240]
a: [3,310,1,2,1,305]
s: 310
i: [1,136,0,4,2,0,0,255] cc9_ordinal1
s: 136
d: [2,312,1,2,1,1,40,136,208]
a: [3,317,1,2,1,312]
s: 317
i: [1,137,0,4,2,0,0,1] cc7_ordinal1
s: 137
i: [1,138,0,4,2,0,0,2] cc3_ordinal1
s: 138
d: [2,313,1,2,1,1,40,136,295]
a: [3,322,1,2,1,313]
s: 322
i: [1,140,0,4,2,0,0,3] cc2_xreal_0
s: 140
d: [2,348,1,2,1,1,40,140,203]
a: [3,355,1,2,1,348]
s: 355
i: [1,141,0,4,2,0,0,2] cc6_ordinal1
s: 141
i: [1,146,0,4,2,0,0,255] cc1_ordinal1
s: 146
d: [2,349,1,2,1,1,40,140,198]
a: [3,360,1,2,1,349]
s: 360
i: [1,147,0,4,2,0,0,255] cc1_ordinal1
s: 147
d: [2,378,1,2,1,1,40,141,203]
a: [3,385,1,2,1,378]
s: 385
i: [1,158,0,4,1,0,0,2] existence_m1_subset_1
s: 158
i: [1,133,0,5,1,0,0,255] t2_boole
s: 133
d: [2,379,1,2,1,1,40,141,198]
a: [3,390,1,2,1,379]
s: 390
i: [1,139,0,5,2,0,0,255] t6_boole
s: 139
d: [2,372,2,2,1,1,40,355,135]
a: [3,376,2,2,1,372]
s: 376
i: [1,149,0,5,2,0,0,1] cc8_ordinal1
s: 149
i: [1,155,0,5,2,0,0,255] t7_boole
s: 155
d: [2,415,2,2,1,1,40,360,135]
a: [3,419,2,2,1,415]
s: 419
i: [1,160,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 160
d: [2,434,2,2,1,1,40,385,147]
a: [3,439,2,2,1,434]
s: 439
a: [3,235,0,5,2,163]
s: 235
d: [2,446,0,5,1,1,22,133,213]
s: 446
d: [2,435,2,2,1,1,40,385,146]
a: [3,444,2,2,1,435]
s: 444
d: [2,459,0,5,2,1,22,139,213]
s: 459
d: [2,447,2,2,1,1,40,390,147]
a: [3,452,2,2,1,447]
s: 452
i: [1,142,0,6,2,0,0,255] d6_card_1
s: 142
i: [1,150,0,6,3,0,0,255] cc2_ordinal1
s: 150
d: [2,448,2,2,1,1,40,390,146]
a: [3,457,2,2,1,448]
s: 457
i: [1,151,0,6,2,0,0,255] t1_subset
s: 151
d: [2,324,1,2,1,1,40,137,208]
a: [3,329,1,2,1,324]
s: 329
i: [1,162,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 162
i: [1,143,0,7,2,0,0,255] redefinition_k7_card_1
s: 143
d: [2,325,1,2,1,1,40,137,295]
a: [3,334,1,2,1,325]
s: 334
i: [1,144,0,7,2,0,0,1] dt_k7_card_1
s: 144
d: [2,336,1,2,1,1,40,138,208]
a: [3,341,1,2,1,336]
s: 341
i: [1,148,0,7,3,0,0,2] cc5_ordinal1
s: 148
i: [1,153,0,7,2,0,0,255] t3_subset
s: 153
d: [2,337,1,2,1,1,40,138,295]
a: [3,346,1,2,1,337]
s: 346
i: [1,154,0,7,3,0,0,1] t8_boole
s: 154
d: [2,350,1,2,1,1,40,140,245]
a: [3,365,1,2,1,350]
s: 365
i: [1,161,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 161
i: [1,134,0,8,3,0,0,255] cc10_ordinal1
s: 134
d: [2,351,1,2,1,1,40,140,275]
a: [3,370,1,2,1,351]
s: 370
i: [1,152,0,8,3,0,0,255] t2_subset
s: 152
d: [2,380,1,2,1,1,40,141,245]
a: [3,395,1,2,1,380]
s: 395
i: [1,157,0,9,3,0,0,255] t5_subset
s: 157
i: [1,156,0,10,3,0,0,255] t4_subset
s: 156
d: [2,381,1,2,1,1,40,141,275]
a: [3,400,1,2,1,381]
s: 400
i: [1,164,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 164
d: [2,604,1,7,3,1,40,164,184]
d: [2,607,1,5,2,2,21,604,376]
d: [2,608,1,3,1,3,21,607,419]
a: [3,612,1,3,1,608]
s: 612
i: [1,165,0,10,2,0,0,1] idempotence_k9_subset_1
s: 165
i: [1,166,0,11,2,0,0,1] dt_k9_subset_1
s: 166
d: [2,402,1,2,1,1,40,146,223]
a: [3,408,1,2,1,402]
s: 408
i: [1,145,0,12,4,0,0,1] t46_nat_1
s: 145
d: [2,403,1,2,1,1,40,146,250]
a: [3,413,1,2,1,403]
s: 413
i: [1,167,0,12,2,0,0,1] redefinition_k9_subset_1
s: 167
i: [1,168,0,13,2,0,0,1] commutativity_k9_subset_1
s: 168
d: [2,629,1,13,2,1,44,167,193]
d: [2,635,1,13,2,1,22,629,161]
a: [3,639,1,5,1,635]
s: 639
d: [2,668,2,0,0,2,52,203,639,144]
e: 668
d: [2,669,2,2,1,1,40,639,144]
d: [2,672,2,0,0,2,21,669,203]
e: 672
r: 639
a: [3,189,0,9,1,126]
s: 189
a: [3,638,1,5,1,635]
s: 638
d: [2,649,1,14,2,1,44,168,193]
a: [3,664,1,5,1,649]
s: 664
d: [2,700,2,0,0,2,52,198,664,144]
e: 700
d: [2,701,2,2,1,1,40,664,144]
d: [2,704,2,0,0,2,21,701,198]
e: 704
r: 664
a: [3,663,1,5,1,649]
s: 663
d: [2,680,2,7,2,1,40,638,156]
s: 680
a: [3,643,1,8,1,635]
s: 643
d: [2,706,2,7,2,1,40,663,156]
s: 706
d: [2,421,1,2,1,1,40,147,223]
a: [3,427,1,2,1,421]
s: 427
d: [2,422,1,2,1,1,40,147,250]
a: [3,432,1,2,1,422]
s: 432
d: [2,540,1,7,3,0,40,134,144]
a: [3,545,1,2,1,540]
s: 545
d: [2,461,1,3,1,0,40,149,158]
a: [3,466,1,3,1,461]
s: 466
d: [2,477,2,2,1,1,40,329,140]
a: [3,483,2,2,1,477]
s: 483
d: [2,507,1,8,3,0,40,148,144]
a: [3,511,1,3,1,507]
s: 511
d: [2,557,1,10,3,0,40,152,144]
a: [3,566,1,3,1,557]
s: 566
d: [2,747,2,0,0,2,52,511,566,138]
e: 747
d: [2,748,2,3,1,1,40,566,138]
d: [2,752,2,0,0,2,21,748,511]
e: 752
r: 566
d: [2,683,2,8,2,1,40,638,152]
a: [3,693,2,5,1,683]
s: 693
a: [3,565,1,3,1,557]
s: 565
d: [2,709,2,8,2,1,40,663,152]
a: [3,719,2,5,1,709]
s: 719
d: [2,558,1,5,2,1,40,152,303]
d: [2,568,1,3,1,2,21,558,228]
a: [3,572,1,3,1,568]
s: 572
d: [2,471,1,5,2,2,44,303,459]
s: 471
d: [2,754,3,5,1,1,40,693,162]
a: [3,760,3,5,1,754]
s: 760
d: [2,506,1,5,2,0,40,148,158]
s: 506
d: [2,763,3,5,1,1,40,719,162]
a: [3,769,3,5,1,763]
s: 769
d: [2,535,1,5,1,1,44,161,446]
s: 535
d: [2,536,1,5,1,1,44,161,446]
s: 536
d: [2,492,2,2,1,1,40,341,147]
a: [3,498,2,2,1,492]
s: 498
d: [2,537,1,5,1,1,44,446,161]
s: 537
d: [2,493,2,2,1,1,40,341,146]
a: [3,503,2,2,1,493]
s: 503
d: [2,538,1,5,1,1,44,446,161]
s: 538
d: [2,486,1,8,3,0,44,142,143]
d: [2,487,1,6,2,0,18,486]
s: 487
r: 511
r: 565
r: 693
r: 719
d: [2,684,2,6,2,1,40,638,148]
a: [3,698,2,3,1,684]
s: 698
r: 698
s: 511
a: [3,697,2,3,1,684]
s: 697
d: [2,897,3,4,2,1,44,697,487]
d: [2,898,3,2,1,2,21,897,385]
d: [2,899,3,0,0,3,21,898,203]
e: 899
r: 697
s: 565
a: [3,514,1,5,2,507]
s: 514
s: 698
d: [2,485,1,8,3,0,44,143,142]
d: [2,488,1,6,2,0,18,485]
s: 488
d: [2,710,2,6,2,1,40,663,148]
a: [3,724,2,3,1,710]
s: 724
r: 724
d: [2,541,1,6,2,0,40,134,158]
s: 541
d: [2,559,1,6,2,0,40,152,158]
s: 559
a: [3,723,2,3,1,710]
s: 723
d: [2,946,3,0,0,3,52,198,723,514]
e: 946
d: [2,947,3,2,1,2,40,723,514]
d: [2,950,3,0,0,3,21,947,198]
e: 950
d: [2,948,3,4,2,1,44,723,487]
d: [2,952,3,2,1,2,21,948,390]
d: [2,953,3,0,0,3,21,952,198]
e: 953
r: 723
d: [2,468,1,7,3,1,44,459,459]
s: 468
s: 724
d: [2,469,1,7,2,1,44,446,459]
s: 469
d: [2,539,1,7,3,0,40,134,153]
s: 539
d: [2,884,3,3,1,1,40,698,147]
a: [3,890,3,3,1,884]
s: 890
a: [3,562,1,7,2,557]
s: 562
d: [2,885,3,3,1,1,40,698,146]
a: [3,895,3,3,1,885]
s: 895
d: [2,588,1,7,2,0,40,157,158]
s: 588
d: [2,516,1,8,3,0,40,153,148]
s: 516
d: [2,910,3,3,1,1,40,724,147]
a: [3,916,3,3,1,910]
s: 916
d: [2,586,1,8,3,0,40,157,153]
s: 586
r: 511
r: 562
r: 565
d: [2,911,3,3,1,1,40,724,146]
a: [3,921,3,3,1,911]
s: 921
s: 566
r: 566
s: 565
d: [2,927,3,8,2,1,40,559,706]
a: [3,943,3,3,1,927]
s: 943
s: 562
d: [2,1015,4,4,2,1,44,943,487]
d: [2,1016,4,2,1,2,21,1015,198]
a: [3,1020,4,2,1,1016]
s: 1020
d: [2,591,1,8,2,0,40,156,158]
s: 591
d: [2,621,1,10,4,1,44,145,446]
d: [2,627,1,8,3,2,21,621,329]
s: 627
d: [2,797,2,6,2,1,44,663,487]
d: [2,828,2,4,1,2,21,797,198]
a: [3,832,2,4,1,828]
s: 832
d: [2,624,1,10,4,1,44,446,145]
d: [2,628,1,8,3,2,21,624,329]
s: 628
d: [2,803,2,6,2,1,44,638,487]
d: [2,859,2,4,1,2,21,803,203]
a: [3,863,2,4,1,859]
s: 863
d: [2,589,1,9,3,0,40,156,153]
s: 589
d: [2,590,1,9,3,0,40,156,144]
s: 590
d: [2,799,3,6,2,1,44,719,487]
d: [2,835,3,4,1,2,21,799,198]
a: [3,839,3,4,1,835]
s: 839
d: [2,556,1,10,3,0,40,152,153]
s: 556
d: [2,806,3,6,2,1,44,693,487]
d: [2,872,3,4,1,2,21,806,203]
a: [3,876,3,4,1,872]
s: 876
d: [2,633,1,14,3,0,44,166,167]
d: [2,634,1,10,2,0,18,633]
s: 634
d: [2,616,1,11,3,0,40,166,134]
s: 616
d: [2,800,4,6,2,1,44,769,487]
d: [2,841,4,4,1,2,21,800,198]
a: [3,845,4,4,1,841]
s: 845
d: [2,653,1,15,3,0,44,166,168]
d: [2,657,1,11,2,0,18,653]
s: 657
d: [2,807,4,6,2,1,44,760,487]
d: [2,878,4,4,1,2,21,807,203]
a: [3,882,4,4,1,878]
s: 882
d: [2,615,1,12,3,0,40,166,157]
s: 615
d: [2,618,1,12,3,0,40,166,148]
s: 618
s: 693
d: [2,622,1,12,4,0,44,145,161]
s: 622
d: [2,1138,2,16,4,1,44,622,643]
a: [3,1151,2,3,1,1138]
s: 1151
d: [2,1162,3,4,2,1,44,1151,487]
d: [2,1163,3,2,1,1,18,1162]
d: [2,1164,3,0,0,2,21,1163,198]
e: 1164
r: 1151
d: [2,623,1,12,4,0,44,145,161]
s: 623
d: [2,625,1,12,4,0,44,161,145]
s: 625
a: [3,1150,2,3,1,1138]
s: 1150
d: [2,626,1,12,4,0,44,161,145]
s: 626
a: [3,1155,2,3,1,1138]
s: 1155
d: [2,1174,3,4,2,1,44,1155,487]
d: [2,1175,3,2,1,1,18,1174]
d: [2,1176,3,0,0,2,21,1175,203]
e: 1176
r: 1155
d: [2,652,1,16,3,0,44,167,168]
d: [2,658,1,12,2,0,18,652]
s: 658
d: [2,648,1,16,3,0,44,168,167]
d: [2,659,1,12,2,0,18,648]
s: 659
a: [3,1154,2,3,1,1138]
s: 1154
d: [2,614,1,13,3,0,40,166,156]
s: 614
d: [2,1167,3,3,1,1,40,1150,140]
a: [3,1172,3,3,1,1167]
s: 1172
d: [2,617,1,14,3,0,40,166,152]
s: 617
d: [2,650,1,16,3,0,44,168,167]
s: 650
d: [2,1207,3,3,1,1,40,1154,140]
a: [3,1212,3,3,1,1207]
s: 1212
d: [2,655,1,16,3,0,44,167,168]
s: 655
d: [2,1224,4,3,1,1,40,1172,135]
a: [3,1229,4,3,1,1224]
s: 1229
s: 719
a: [3,1159,2,5,1,1138]
s: 1159
r: 1159
d: [2,1245,4,3,1,1,40,1212,135]
a: [3,1250,4,3,1,1245]
s: 1250
d: [2,963,2,8,2,2,44,469,643]
a: [3,974,2,5,1,963]
s: 974
r: 189
r: 193
r: 643
r: 663
r: 706
r: 719
r: 724
r: 769
r: 916
r: 921
r: 943
r: 1150
r: 1172
r: 1229
r: 974
s: 724
s: 189
s: 193
s: 1150
s: 663
s: 916
s: 643
s: 719
s: 921
a: [3,973,2,5,1,963]
s: 973
s: 943
s: 1159
d: [2,804,2,8,2,1,44,643,487]
d: [2,865,2,6,1,2,21,804,203]
a: [3,869,2,6,1,865]
s: 869
s: 1172
s: 706
s: 1229
d: [2,796,2,9,2,1,44,643,487]
d: [2,822,2,7,1,2,21,796,198]
a: [3,826,2,7,1,822]
s: 826
d: [2,801,2,9,2,1,44,189,487]
d: [2,847,2,7,1,2,21,801,203]
a: [3,851,2,7,1,847]
s: 851
d: [2,1261,3,6,2,1,44,1159,487]
d: [2,1269,3,4,1,2,21,1261,198]
a: [3,1273,3,4,1,1269]
s: 1273
d: [2,1512,4,5,2,1,44,1273,487]
d: [2,1520,4,2,1,2,21,1512,612]
d: [2,1521,4,0,0,3,21,1520,203]
e: 1521
d: [2,654,1,14,2,1,44,193,168]
d: [2,1327,1,13,2,1,44,167,189]
d: [2,1386,2,8,2,1,44,643,487]
d: [2,1393,2,6,1,2,21,1386,203]
d: [2,1387,2,9,2,1,44,643,487]
d: [2,1395,2,7,1,2,21,1387,198]
d: [2,1390,2,16,4,1,44,622,643]
d: [2,1401,2,13,3,2,21,1390,1154]
d: [2,1402,2,10,2,3,21,1401,1150]
d: [2,1419,3,6,2,1,44,1159,487]
d: [2,1423,3,4,1,2,21,1419,198]
d: [2,1430,3,13,4,1,44,622,869]
d: [2,1446,3,11,3,2,21,1430,203]
d: [2,1447,3,8,2,3,21,1446,1150]
d: [2,1463,3,13,4,1,44,622,826]
d: [2,1479,3,10,3,2,21,1463,1154]
d: [2,1480,3,8,2,3,21,1479,198]
f: 1521,1480,1447,1423,1402,1395,1393,1327,1176,1164,704,654,608,415,372,349,348,124,125,126,127
% Refutation found. Thanks to Tanya!
% SZS status Theorem for algseq_1__t7_algseq_1
% SZS output start Proof for algseq_1__t7_algseq_1
1. ! [X0] : (v7_ordinal1(X0) => ! [X1] : (v7_ordinal1(X1) => ((k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X1),k7_card_1(X0)) | k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X0),k7_card_1(X1))) => r1_xxreal_0(X0,X1)))) [input t7_algseq_1]
2. ~! [X0] : (v7_ordinal1(X0) => ! [X1] : (v7_ordinal1(X1) => ((k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X1),k7_card_1(X0)) | k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X0),k7_card_1(X1))) => r1_xxreal_0(X0,X1)))) [negated conjecture 1]
25. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
38. ! [X0] : (v7_ordinal1(X0) => k6_card_1(X0) = X0) [input d6_card_1]
39. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
41. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => k9_subset_1(X0,X1,X2) = k9_subset_1(X0,X2,X1)) [input commutativity_k9_subset_1]
44. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
45. ! [X0] : (v7_ordinal1(X0) => k7_card_1(X0) = k6_card_1(X0)) [input redefinition_k7_card_1]
46. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(X0)) => k9_subset_1(X0,X1,X2) = k3_xboole_0(X1,X2)) [input redefinition_k9_subset_1]
49. ! [X0] : (v7_ordinal1(X0) => m1_subset_1(k7_card_1(X0),k1_zfmisc_1(k4_ordinal1))) [input dt_k7_card_1]
51. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
55. ! [X0] : (v7_ordinal1(X0) => ! [X1] : (v7_ordinal1(X1) => (r1_xxreal_0(X0,X1) => k3_xboole_0(X0,X1) = X0))) [input t46_nat_1]
69. ? [X0] : (? [X1] : ((~r1_xxreal_0(X0,X1) & (k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X1),k7_card_1(X0)) | k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X0),k7_card_1(X1)))) & v7_ordinal1(X1)) & v7_ordinal1(X0)) [ennf transformation 2]
70. ? [X0] : (? [X1] : (~r1_xxreal_0(X0,X1) & (k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X1),k7_card_1(X0)) | k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X0),k7_card_1(X1))) & v7_ordinal1(X1)) & v7_ordinal1(X0)) [flattening 69]
72. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 25]
77. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 51]
79. ! [X0] : (k6_card_1(X0) = X0 | ~v7_ordinal1(X0)) [ennf transformation 38]
80. ! [X0] : (k7_card_1(X0) = k6_card_1(X0) | ~v7_ordinal1(X0)) [ennf transformation 45]
81. ! [X0] : (m1_subset_1(k7_card_1(X0),k1_zfmisc_1(k4_ordinal1)) | ~v7_ordinal1(X0)) [ennf transformation 49]
82. ! [X0] : (! [X1] : ((k3_xboole_0(X0,X1) = X0 | ~r1_xxreal_0(X0,X1)) | ~v7_ordinal1(X1)) | ~v7_ordinal1(X0)) [ennf transformation 55]
83. ! [X0] : (! [X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_xxreal_0(X0,X1) | ~v7_ordinal1(X1)) | ~v7_ordinal1(X0)) [flattening 82]
101. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 44]
102. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 101]
105. ! [X0,X1,X2] : (k9_subset_1(X0,X1,X2) = k3_xboole_0(X1,X2) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) [ennf transformation 46]
106. ! [X0,X1,X2] : (k9_subset_1(X0,X1,X2) = k9_subset_1(X0,X2,X1) | ~m1_subset_1(X2,k1_zfmisc_1(X0))) [ennf transformation 41]
107. ? [X0] : (? [X1] : (~r1_xxreal_0(X0,X1) & (k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X1),k7_card_1(X0)) | k7_card_1(X0) = k9_subset_1(k4_ordinal1,k7_card_1(X0),k7_card_1(X1))) & v7_ordinal1(X1)) & v7_ordinal1(X0)) => (? [X1] : (~r1_xxreal_0(sK0,X1) & (k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(X1),k7_card_1(sK0)) | k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK0),k7_card_1(X1))) & v7_ordinal1(X1)) & v7_ordinal1(sK0)) [choice axiom]
108. ? [X1] : (~r1_xxreal_0(sK0,X1) & (k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(X1),k7_card_1(sK0)) | k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK0),k7_card_1(X1))) & v7_ordinal1(X1)) => (~r1_xxreal_0(sK0,sK1) & (k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK1),k7_card_1(sK0)) | k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK0),k7_card_1(sK1))) & v7_ordinal1(sK1)) [choice axiom]
109. (~r1_xxreal_0(sK0,sK1) & (k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK1),k7_card_1(sK0)) | k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK0),k7_card_1(sK1))) & v7_ordinal1(sK1)) & v7_ordinal1(sK0) [skolemisation 70,108,107]
124. v7_ordinal1(sK0) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
125. v7_ordinal1(sK1) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
126. k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK1),k7_card_1(sK0)) | k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK0),k7_card_1(sK1)) [cnf transformation 109] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
127. ~r1_xxreal_0(sK0,sK1) [cnf transformation 109] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
135. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 72] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
140. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 77] {a:0,w:4,wCS:1600,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
142. k6_card_1(X0) = X0 | ~v7_ordinal1(X0) [cnf transformation 79] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
143. k7_card_1(X0) = k6_card_1(X0) | ~v7_ordinal1(X0) [cnf transformation 80] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
144. m1_subset_1(k7_card_1(X0),k1_zfmisc_1(k4_ordinal1)) | ~v7_ordinal1(X0) [cnf transformation 81] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
145. k3_xboole_0(X0,X1) = X0 | ~r1_xxreal_0(X0,X1) | ~v7_ordinal1(X1) | ~v7_ordinal1(X0) [cnf transformation 83] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
161. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 39] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
164. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 102] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
167. k9_subset_1(X0,X1,X2) = k3_xboole_0(X1,X2) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) [cnf transformation 105] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
168. k9_subset_1(X0,X1,X2) = k9_subset_1(X0,X2,X1) | ~m1_subset_1(X2,k1_zfmisc_1(X0)) [cnf transformation 106] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
182. 1 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
184. ~r1_xxreal_0(sK0,sK1) <- (~1) [avatar component clause 182] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
185. ~1 [avatar split clause 127,182]
187. 2 <=> k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK0),k7_card_1(sK1)) [avatar definition]
189. k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK0),k7_card_1(sK1)) <- (2) [avatar component clause 187] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
191. 3 <=> k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK1),k7_card_1(sK0)) [avatar definition]
193. k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK1),k7_card_1(sK0)) <- (3) [avatar component clause 191] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
194. 2 | 3 [avatar split clause 126,191,187]
196. 4 <=> v7_ordinal1(sK1) [avatar definition]
198. v7_ordinal1(sK1) <- (4) [avatar component clause 196] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
199. 4 [avatar split clause 125,196]
201. 5 <=> v7_ordinal1(sK0) [avatar definition]
203. v7_ordinal1(sK0) <- (5) [avatar component clause 201] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
204. 5 [avatar split clause 124,201]
348. v1_xreal_0(sK0) <- (5) [resolution 140,203] {a:1,w:2,wCS:200,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
349. v1_xreal_0(sK1) <- (4) [resolution 140,198] {a:1,w:2,wCS:200,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
353. 33 <=> v1_xreal_0(sK0) [avatar definition]
355. v1_xreal_0(sK0) <- (33) [avatar component clause 353] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
356. 33 | ~5 [avatar split clause 348,201,353]
358. 34 <=> v1_xreal_0(sK1) [avatar definition]
360. v1_xreal_0(sK1) <- (34) [avatar component clause 358] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
361. 34 | ~4 [avatar split clause 349,196,358]
372. v1_xxreal_0(sK0) <- (33) [resolution 355,135] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
374. 37 <=> v1_xxreal_0(sK0) [avatar definition]
376. v1_xxreal_0(sK0) <- (37) [avatar component clause 374] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
377. 37 | ~33 [avatar split clause 372,353,374]
415. v1_xxreal_0(sK1) <- (34) [resolution 360,135] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
417. 44 <=> v1_xxreal_0(sK1) [avatar definition]
419. v1_xxreal_0(sK1) <- (44) [avatar component clause 417] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
420. 44 | ~34 [avatar split clause 415,358,417]
486. k7_card_1(X0) = X0 | ~v7_ordinal1(X0) | ~v7_ordinal1(X0) [superposition 142,143] {a:1,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
487. k7_card_1(X0) = X0 | ~v7_ordinal1(X0) [duplicate literal removal 486] {a:1,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
604. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~1) [resolution 164,184] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
607. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) <- (~1, 37) [subsumption resolution 604,376] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
608. r1_xxreal_0(sK1,sK0) <- (~1, 37, 44) [subsumption resolution 607,419] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
610. 70 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
612. r1_xxreal_0(sK1,sK0) <- (70) [avatar component clause 610] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
613. 70 | 1 | ~37 | ~44 [avatar split clause 608,417,374,182,610]
622. k3_xboole_0(X5,X4) = X4 | ~r1_xxreal_0(X4,X5) | ~v7_ordinal1(X5) | ~v7_ordinal1(X4) [superposition 145,161] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
641. 72 <=> k7_card_1(sK0) = k3_xboole_0(k7_card_1(sK0),k7_card_1(sK1)) [avatar definition]
643. k7_card_1(sK0) = k3_xboole_0(k7_card_1(sK0),k7_card_1(sK1)) <- (72) [avatar component clause 641] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
654. k7_card_1(sK0) = k9_subset_1(k4_ordinal1,k7_card_1(sK0),k7_card_1(sK1)) | ~m1_subset_1(k7_card_1(sK1),k1_zfmisc_1(k4_ordinal1)) <- (3) [superposition 193,168] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
662. 73 <=> m1_subset_1(k7_card_1(sK1),k1_zfmisc_1(k4_ordinal1)) [avatar definition]
664. ~m1_subset_1(k7_card_1(sK1),k1_zfmisc_1(k4_ordinal1)) <- (~73) [avatar component clause 662] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
667. ~73 | 2 | ~3 [avatar split clause 654,191,187,662]
701. ~v7_ordinal1(sK1) <- (~73) [resolution 664,144] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
704. $false <- (4, ~73) [subsumption resolution 701,198] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
705. ~4 | 73 [avatar contradiction clause 704]
824. 88 <=> k7_card_1(sK0) = k3_xboole_0(k7_card_1(sK0),sK1) [avatar definition]
826. k7_card_1(sK0) = k3_xboole_0(k7_card_1(sK0),sK1) <- (88) [avatar component clause 824] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
867. 95 <=> sK0 = k3_xboole_0(sK0,k7_card_1(sK1)) [avatar definition]
869. sK0 = k3_xboole_0(sK0,k7_card_1(sK1)) <- (95) [avatar component clause 867] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
972. 106 <=> k7_card_1(sK0) = k7_card_1(sK1) [avatar definition]
1149. 113 <=> v7_ordinal1(k7_card_1(sK1)) [avatar definition]
1150. v7_ordinal1(k7_card_1(sK1)) <- (113) [avatar component clause 1149] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1151. ~v7_ordinal1(k7_card_1(sK1)) <- (~113) [avatar component clause 1149] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1153. 114 <=> v7_ordinal1(k7_card_1(sK0)) [avatar definition]
1154. v7_ordinal1(k7_card_1(sK0)) <- (114) [avatar component clause 1153] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1155. ~v7_ordinal1(k7_card_1(sK0)) <- (~114) [avatar component clause 1153] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1157. 115 <=> r1_xxreal_0(k7_card_1(sK1),k7_card_1(sK0)) [avatar definition]
1159. ~r1_xxreal_0(k7_card_1(sK1),k7_card_1(sK0)) <- (~115) [avatar component clause 1157] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1162. ~v7_ordinal1(sK1) | ~v7_ordinal1(sK1) <- (~113) [superposition 1151,487] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1163. ~v7_ordinal1(sK1) <- (~113) [duplicate literal removal 1162] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1164. $false <- (4, ~113) [subsumption resolution 1163,198] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1165. ~4 | 113 [avatar contradiction clause 1164]
1174. ~v7_ordinal1(sK0) | ~v7_ordinal1(sK0) <- (~114) [superposition 1155,487] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1175. ~v7_ordinal1(sK0) <- (~114) [duplicate literal removal 1174] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1176. $false <- (5, ~114) [subsumption resolution 1175,203] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1177. ~5 | 114 [avatar contradiction clause 1176]
1271. 122 <=> r1_xxreal_0(sK1,k7_card_1(sK0)) [avatar definition]
1273. ~r1_xxreal_0(sK1,k7_card_1(sK0)) <- (~122) [avatar component clause 1271] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1277. 123 <=> r1_xxreal_0(k7_card_1(sK1),sK0) [avatar definition]
1311. 125 <=> k7_card_1(sK0) = sK1 [avatar definition]
1327. k7_card_1(sK0) = k3_xboole_0(k7_card_1(sK0),k7_card_1(sK1)) | ~m1_subset_1(k7_card_1(sK1),k1_zfmisc_1(k4_ordinal1)) <- (2) [superposition 167,189] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1342. ~73 | 72 | ~2 [avatar split clause 1327,187,641,662]
1386. sK0 = k3_xboole_0(sK0,k7_card_1(sK1)) | ~v7_ordinal1(sK0) <- (72) [superposition 643,487] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1387. k7_card_1(sK0) = k3_xboole_0(k7_card_1(sK0),sK1) | ~v7_ordinal1(sK1) <- (72) [superposition 643,487] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1390. k7_card_1(sK0) = k7_card_1(sK1) | ~r1_xxreal_0(k7_card_1(sK1),k7_card_1(sK0)) | ~v7_ordinal1(k7_card_1(sK0)) | ~v7_ordinal1(k7_card_1(sK1)) <- (72) [superposition 622,643] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1393. sK0 = k3_xboole_0(sK0,k7_card_1(sK1)) <- (5, 72) [subsumption resolution 1386,203] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1394. 95 | ~5 | ~72 [avatar split clause 1393,641,201,867]
1395. k7_card_1(sK0) = k3_xboole_0(k7_card_1(sK0),sK1) <- (4, 72) [subsumption resolution 1387,198] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1396. 88 | ~4 | ~72 [avatar split clause 1395,641,196,824]
1401. k7_card_1(sK0) = k7_card_1(sK1) | ~r1_xxreal_0(k7_card_1(sK1),k7_card_1(sK0)) | ~v7_ordinal1(k7_card_1(sK1)) <- (72, 114) [subsumption resolution 1390,1154] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1402. k7_card_1(sK0) = k7_card_1(sK1) | ~r1_xxreal_0(k7_card_1(sK1),k7_card_1(sK0)) <- (72, 113, 114) [subsumption resolution 1401,1150] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1403. ~115 | 106 | ~72 | ~113 | ~114 [avatar split clause 1402,1153,1149,641,972,1157]
1419. ~r1_xxreal_0(sK1,k7_card_1(sK0)) | ~v7_ordinal1(sK1) <- (~115) [superposition 1159,487] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1423. ~r1_xxreal_0(sK1,k7_card_1(sK0)) <- (4, ~115) [subsumption resolution 1419,198] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1424. ~122 | ~4 | 115 [avatar split clause 1423,1157,196,1271]
1430. sK0 = k7_card_1(sK1) | ~r1_xxreal_0(k7_card_1(sK1),sK0) | ~v7_ordinal1(sK0) | ~v7_ordinal1(k7_card_1(sK1)) <- (95) [superposition 622,869] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1442. 127 <=> sK0 = k7_card_1(sK1) [avatar definition]
1446. sK0 = k7_card_1(sK1) | ~r1_xxreal_0(k7_card_1(sK1),sK0) | ~v7_ordinal1(k7_card_1(sK1)) <- (5, 95) [subsumption resolution 1430,203] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1447. sK0 = k7_card_1(sK1) | ~r1_xxreal_0(k7_card_1(sK1),sK0) <- (5, 95, 113) [subsumption resolution 1446,1150] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1448. ~123 | 127 | ~5 | ~95 | ~113 [avatar split clause 1447,1149,867,201,1442,1277]
1463. k7_card_1(sK0) = sK1 | ~r1_xxreal_0(sK1,k7_card_1(sK0)) | ~v7_ordinal1(k7_card_1(sK0)) | ~v7_ordinal1(sK1) <- (88) [superposition 622,826] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1479. k7_card_1(sK0) = sK1 | ~r1_xxreal_0(sK1,k7_card_1(sK0)) | ~v7_ordinal1(sK1) <- (88, 114) [subsumption resolution 1463,1154] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1480. k7_card_1(sK0) = sK1 | ~r1_xxreal_0(sK1,k7_card_1(sK0)) <- (4, 88, 114) [subsumption resolution 1479,198] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
1481. ~122 | 125 | ~4 | ~88 | ~114 [avatar split clause 1480,1153,824,196,1311,1271]
1512. ~r1_xxreal_0(sK1,sK0) | ~v7_ordinal1(sK0) <- (~122) [superposition 1273,487] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
1520. ~v7_ordinal1(sK0) <- (70, ~122) [subsumption resolution 1512,612] {a:4,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
1521. $false <- (5, 70, ~122) [subsumption resolution 1520,203] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
1522. ~5 | ~70 | 122 [avatar contradiction clause 1521]
1523. k7_card_1(sK0) != k7_card_1(sK1) | k7_card_1(sK0) != sK1 | r1_xxreal_0(k7_card_1(sK1),sK0) | ~r1_xxreal_0(sK1,sK0) [theory tautology sat conflict] {a:0,w:16,wCS:6400,sine:255,thAx:0,allAx:1,thDist:-1}
1524. sK0 != k7_card_1(sK1) | k7_card_1(sK0) != k7_card_1(sK1) | k7_card_1(sK0) != sK1 | r1_xxreal_0(sK0,sK1) | ~r1_xxreal_0(k7_card_1(sK1),k7_card_1(sK0)) [theory tautology sat conflict] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:1,thDist:-1}
1525. $false [avatar sat refutation 185,194,199,204,356,361,377,420,613,667,705,1165,1177,1342,1394,1396,1403,1424,1448,1481,1522,1523,1524] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1485,thDist:-1485}
% SZS output end Proof for algseq_1__t7_algseq_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 54
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 57
% Generated clauses: 841
% Active clauses: 186
% Passive clauses: 326
% Final active clauses: 152
% Final passive clauses: 116
% Main loop iterations started: 202
% 
% >>> Simplifying Inferences
% Duplicate literals: 32
% Fw subsumption resolutions: 82
% Fw demodulations: 5
% Bw demodulations: 14
% 
% >>> Deletion Inferences
% Simple tautologies: 19
% Equational tautologies: 1
% Forward subsumptions: 187
% 
% >>> Generating Inferences
% Binary resolution: 227
% Unit resulting resolution: 4
% Factoring: 1
% Forward superposition: 154
% Backward superposition: 105
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 43
% Split components: 90
% Sat splitting refutations: 14
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 207
% SAT solver unit clauses: 21
% SAT solver binary clauses: 91
% 
% Memory used [KB]: 5884
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% resolution: 0.001 s
% other: 0.006 s ( own 0.004 s ) 

