i: [1,62,0,5,1,1,0,0] t6_xtuple_0
a: [3,106,0,5,1,62]
s: 106
i: [1,61,0,5,1,1,0,0] t6_xtuple_0
a: [3,111,0,5,1,61]
s: 111
i: [1,63,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,116,0,2,1,63]
s: 116
i: [1,64,0,2,1,0,0,255] fc1_xboole_0
a: [3,121,0,2,1,64]
s: 121
i: [1,90,0,2,1,0,0,2] rc2_xboole_0
a: [3,131,0,2,1,90]
s: 131
i: [1,91,0,2,1,0,0,2] rc1_xboole_0
a: [3,136,0,2,1,91]
s: 136
i: [1,66,0,3,1,0,0,2] fc2_xboole_0
s: 66
i: [1,65,0,3,1,0,0,255] d2_xboole_0
a: [3,126,0,3,1,65]
s: 126
r: 116
i: [1,87,0,4,1,0,0,2] fc3_xboole_0
s: 87
i: [1,93,0,4,1,0,0,2] d1_tarski
s: 93
i: [1,67,0,5,2,0,0,255] t6_boole
s: 67
i: [1,80,0,5,2,0,0,1] t7_boole
s: 80
i: [1,99,0,5,1,0,0,2] d2_tarski
s: 99
d: [2,141,1,2,1,1,40,80,111]
a: [3,146,1,2,1,141]
s: 146
i: [1,101,0,5,1,0,0,2] d2_tarski
s: 101
i: [1,89,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 89
d: [2,150,1,5,1,1,40,89,111]
a: [3,157,1,5,1,150]
s: 157
i: [1,79,0,7,3,0,0,2] t8_boole
s: 79
d: [2,151,1,4,1,0,40,89,93]
s: 151
i: [1,88,0,7,1,0,0,2] commutativity_k2_tarski
s: 88
i: [1,94,0,7,2,0,0,2] d1_tarski
s: 94
d: [2,152,1,5,1,0,40,89,99]
s: 152
i: [1,96,0,9,2,0,0,1] d4_tarski
s: 96
d: [2,153,1,5,1,0,40,89,101]
s: 153
i: [1,97,0,9,2,0,0,1] d4_tarski
s: 97
i: [1,68,0,10,1,0,0,1] d5_tarski
s: 68
d: [2,180,1,4,1,0,44,87,68]
s: 180
i: [1,95,0,10,3,0,0,1] d4_tarski
s: 95
d: [2,192,1,7,2,1,40,95,106]
s: 192
i: [1,102,0,11,3,0,0,2] d2_tarski
s: 102
i: [1,71,0,14,3,0,0,2] d1_tarski
s: 71
d: [2,193,2,9,3,1,40,192,95]
s: 193
i: [1,72,0,14,3,0,0,2] d1_tarski
s: 72
d: [2,209,3,7,2,1,40,193,93]
s: 209
i: [1,77,0,14,3,0,0,1] d4_tarski
s: 77
i: [1,76,0,16,3,0,0,1] d4_tarski
s: 76
d: [2,214,4,5,1,1,40,209,93]
a: [3,219,4,5,1,214]
s: 219
i: [1,78,0,17,4,0,0,1] d4_tarski
s: 78
d: [2,212,4,6,1,1,40,209,99]
s: 212
i: [1,85,0,17,3,0,0,2] d2_tarski
s: 85
i: [1,86,0,17,3,0,0,2] d2_tarski
s: 86
d: [2,237,5,5,1,1,44,212,68]
s: 237
d: [2,239,6,0,0,2,40,237,111]
e: 239
d: [2,238,6,0,0,2,52,111,237]
e: 238
f: 238,61,62
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t6_xtuple_0
% SZS output start Proof for xtuple_0__t6_xtuple_0
1. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(k4_tarski(X0,X1),X2) => r2_hidden(X0,k3_tarski(k3_tarski(X2)))) [input t6_xtuple_0]
2. ~! [X0] : ! [X1] : ! [X2] : (r2_hidden(k4_tarski(X0,X1),X2) => r2_hidden(X0,k3_tarski(k3_tarski(X2)))) [negated conjecture 1]
5. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
6. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
8. ! [X0] : ! [X1] : (k3_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)))) [input d4_tarski]
9. ! [X0] : ! [X1] : k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [input d5_tarski]
26. ~! [X0,X1,X2] : (r2_hidden(k4_tarski(X0,X1),X2) => r2_hidden(X0,k3_tarski(k3_tarski(X2)))) [flattening 2]
27. ! [X0,X1] : k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [flattening 9]
28. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 5]
29. ! [X0,X1] : (k3_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)))) [flattening 8]
32. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 6]
35. ? [X0,X1,X2] : (~r2_hidden(X0,k3_tarski(k3_tarski(X2))) & r2_hidden(k4_tarski(X0,X1),X2)) [ennf transformation 26]
40. ? [X0,X1,X2] : (~r2_hidden(X0,k3_tarski(k3_tarski(X2))) & r2_hidden(k4_tarski(X0,X1),X2)) => (~r2_hidden(sK0,k3_tarski(k3_tarski(sK2))) & r2_hidden(k4_tarski(sK0,sK1),sK2)) [choice axiom]
41. ~r2_hidden(sK0,k3_tarski(k3_tarski(sK2))) & r2_hidden(k4_tarski(sK0,sK1),sK2) [skolemisation 35,40]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 28]
43. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 42]
44. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
45. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 43,44]
46. ! [X0,X1] : ((k3_tarski(X0) = X1 | ? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3))) & (? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)) | ~r2_hidden(X2,X1))) | k3_tarski(X0) != X1)) [nnf transformation 29]
47. ! [X0,X1] : ((k3_tarski(X0) = X1 | ? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(X2,X4)) | r2_hidden(X2,X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (~r2_hidden(X6,X0) | ~r2_hidden(X5,X6))) & (? [X7] : (r2_hidden(X7,X0) & r2_hidden(X5,X7)) | ~r2_hidden(X5,X1))) | k3_tarski(X0) != X1)) [rectify 46]
48. ! [X1,X0] : (? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(X2,X4)) | r2_hidden(X2,X1))) => ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(sK4(X0,X1),X3)) | ~r2_hidden(sK4(X0,X1),X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(sK4(X0,X1),X4)) | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
49. ! [X1,X0] : (? [X4] : (r2_hidden(X4,X0) & r2_hidden(sK4(X0,X1),X4)) => (r2_hidden(sK5(X0,X1),X0) & r2_hidden(sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
50. ! [X5,X0] : (? [X7] : (r2_hidden(X7,X0) & r2_hidden(X5,X7)) => (r2_hidden(sK6(X0,X5),X0) & r2_hidden(X5,sK6(X0,X5)))) [choice axiom]
51. ! [X0,X1] : ((k3_tarski(X0) = X1 | ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(sK4(X0,X1),X3)) | ~r2_hidden(sK4(X0,X1),X1)) & ((r2_hidden(sK5(X0,X1),X0) & r2_hidden(sK4(X0,X1),sK5(X0,X1))) | r2_hidden(sK4(X0,X1),X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (~r2_hidden(X6,X0) | ~r2_hidden(X5,X6))) & ((r2_hidden(sK6(X0,X5),X0) & r2_hidden(X5,sK6(X0,X5))) | ~r2_hidden(X5,X1))) | k3_tarski(X0) != X1)) [skolemisation 47,50,49,48]
52. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 32]
53. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 52]
54. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 53]
55. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK7(X0,X1,X2) != X1 & sK7(X0,X1,X2) != X0) | ~r2_hidden(sK7(X0,X1,X2),X2)) & (sK7(X0,X1,X2) = X1 | sK7(X0,X1,X2) = X0 | r2_hidden(sK7(X0,X1,X2),X2)))) [choice axiom]
56. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK7(X0,X1,X2) != X1 & sK7(X0,X1,X2) != X0) | ~r2_hidden(sK7(X0,X1,X2),X2)) & (sK7(X0,X1,X2) = X1 | sK7(X0,X1,X2) = X0 | r2_hidden(sK7(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 54,55]
61. r2_hidden(k4_tarski(sK0,sK1),sK2) [cnf transformation 41] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
62. ~r2_hidden(sK0,k3_tarski(k3_tarski(sK2))) [cnf transformation 41] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
68. k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [cnf transformation 27] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 45] {a:0,w:10,wCS:4000,sine:2,thAx:0,allAx:2,thDist:-2}
75. r2_hidden(X5,X1) | ~r2_hidden(X6,X0) | ~r2_hidden(X5,X6) | k3_tarski(X0) != X1 [cnf transformation 51] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:4,thDist:-4}
83. r2_hidden(X4,X2) | X1 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 56] {a:0,w:11,wCS:4400,sine:2,thAx:0,allAx:2,thDist:-2}
92. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 70] {a:0,w:7,wCS:2800,sine:2,thAx:0,allAx:2,thDist:-2}
93. r2_hidden(X3,k1_tarski(X3)) [equality resolution 92] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
95. r2_hidden(X5,k3_tarski(X0)) | ~r2_hidden(X6,X0) | ~r2_hidden(X5,X6) [equality resolution 75] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
98. r2_hidden(X4,X2) | k2_tarski(X0,X4) != X2 [equality resolution 83] {a:0,w:8,wCS:3200,sine:2,thAx:0,allAx:2,thDist:-2}
99. r2_hidden(X4,k2_tarski(X0,X4)) [equality resolution 98] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
104. 1 <=> r2_hidden(sK0,k3_tarski(k3_tarski(sK2))) [avatar definition]
106. ~r2_hidden(sK0,k3_tarski(k3_tarski(sK2))) <- (~1) [avatar component clause 104] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
107. ~1 [avatar split clause 62,104]
109. 2 <=> r2_hidden(k4_tarski(sK0,sK1),sK2) [avatar definition]
111. r2_hidden(k4_tarski(sK0,sK1),sK2) <- (2) [avatar component clause 109] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
112. 2 [avatar split clause 61,109]
192. ~r2_hidden(X14,k3_tarski(sK2)) | ~r2_hidden(sK0,X14) <- (~1) [resolution 95,106] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
193. ~r2_hidden(sK0,X0) | ~r2_hidden(X1,sK2) | ~r2_hidden(X0,X1) <- (~1) [resolution 192,95] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
209. ~r2_hidden(k1_tarski(sK0),X7) | ~r2_hidden(X7,sK2) <- (~1) [resolution 193,93] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
212. ~r2_hidden(k2_tarski(X2,k1_tarski(sK0)),sK2) <- (~1) [resolution 209,99] {a:4,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
237. ~r2_hidden(k4_tarski(sK0,X2),sK2) <- (~1) [superposition 212,68] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
238. $false <- (~1, 2) [unit resulting resolution 111,237] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:17,thDist:-17}
241. 1 | ~2 [avatar contradiction clause 238]
242. $false [avatar sat refutation 107,112,241] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:78,thDist:-78}
% SZS output end Proof for xtuple_0__t6_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 31
% Generated clauses: 128
% Active clauses: 42
% Passive clauses: 96
% Final active clauses: 41
% Final passive clauses: 54
% Main loop iterations started: 42
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 3
% Forward subsumptions: 14
% 
% >>> Generating Inferences
% Binary resolution: 54
% Unit resulting resolution: 1
% Forward superposition: 14
% Backward superposition: 14
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 7
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

