i: [1,103,0,6,2,1,0,0] l186_relat_1
a: [3,147,0,3,1,103]
s: 147
i: [1,104,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,156,0,2,1,104]
s: 156
i: [1,105,0,2,1,0,0,1] fc1_xboole_0
a: [3,161,0,2,1,105]
s: 161
i: [1,137,0,2,1,0,0,3] rc2_zfmisc_1
a: [3,171,0,2,1,137]
s: 171
i: [1,138,0,2,1,0,0,2] rc2_xboole_0
a: [3,176,0,2,1,138]
s: 176
i: [1,139,0,2,1,0,0,2] rc1_xboole_0
a: [3,181,0,2,1,139]
s: 181
i: [1,141,0,2,1,0,0,2] rc1_relat_1
a: [3,186,0,2,1,141]
s: 186
i: [1,140,0,2,1,0,0,2] rc1_relat_1
a: [3,191,0,2,1,140]
s: 191
i: [1,143,0,2,1,0,0,3] rc1_zfmisc_1
a: [3,196,0,2,1,143]
s: 196
i: [1,142,0,2,1,0,0,3] rc1_zfmisc_1
a: [3,201,0,2,1,142]
s: 201
i: [1,107,0,3,1,0,0,1] t2_xboole_1
s: 107
i: [1,135,0,3,1,0,0,2] reflexivity_r1_tarski
s: 135
i: [1,106,0,3,1,0,0,255] d2_xboole_0
a: [3,166,0,3,1,106]
s: 166
r: 156
i: [1,108,0,4,2,0,0,3] cc2_zfmisc_1
s: 108
i: [1,110,0,4,2,0,0,2] cc1_relat_1
s: 110
i: [1,111,0,4,2,0,0,3] cc1_zfmisc_1
s: 111
d: [2,204,1,2,1,1,40,108,181]
a: [3,209,1,2,1,204]
s: 209
i: [1,134,0,4,1,0,0,255] existence_m1_subset_1
s: 134
d: [2,205,1,2,1,1,40,108,161]
a: [3,214,1,2,1,205]
s: 214
i: [1,112,0,5,2,0,0,255] t6_boole
s: 112
i: [1,113,0,5,2,0,0,2] fc11_relat_1
s: 113
d: [2,232,1,5,2,1,44,147,112]
s: 232
i: [1,114,0,5,2,0,0,2] fc10_relat_1
s: 114
d: [2,216,1,2,1,1,40,110,181]
a: [3,221,1,2,1,216]
s: 221
i: [1,115,0,5,2,0,0,2] fc5_xtuple_0
s: 115
i: [1,116,0,5,2,0,0,2] fc4_xtuple_0
s: 116
d: [2,217,1,2,1,1,40,110,161]
a: [3,226,1,2,1,217]
s: 226
i: [1,131,0,5,2,0,0,255] t7_boole
s: 131
d: [2,230,1,5,2,0,44,107,112]
s: 230
i: [1,126,0,6,2,0,0,255] t1_subset
s: 126
i: [1,136,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 136
d: [2,233,1,5,2,0,40,113,110]
s: 233
i: [1,117,0,7,3,0,0,2] fc9_relat_1
s: 117
d: [2,234,1,5,2,0,40,113,108]
s: 234
i: [1,118,0,7,3,0,0,2] fc8_relat_1
s: 118
i: [1,119,0,7,3,0,0,2] fc25_relat_1
s: 119
d: [2,235,1,5,2,0,40,114,110]
s: 235
i: [1,120,0,7,3,0,0,2] fc24_relat_1
s: 120
d: [2,236,1,5,2,0,40,114,108]
s: 236
i: [1,128,0,7,2,0,0,255] t3_subset
s: 128
i: [1,129,0,7,2,0,0,255] t3_subset
s: 129
d: [2,240,1,5,1,0,40,128,134]
s: 240
i: [1,130,0,7,3,0,0,255] t8_boole
s: 130
d: [2,228,1,7,3,0,44,112,112]
s: 228
i: [1,109,0,8,3,0,0,255] cc2_relat_1
s: 109
i: [1,127,0,8,3,0,0,255] t2_subset
s: 127
d: [2,242,1,6,2,0,40,109,134]
s: 242
i: [1,122,0,9,3,0,0,1] d18_relat_1
s: 122
d: [2,244,1,6,2,0,40,127,134]
s: 244
i: [1,123,0,9,3,0,0,1] d18_relat_1
s: 123
i: [1,124,0,9,3,0,0,1] d19_relat_1
s: 124
d: [2,249,1,6,2,0,40,123,135]
s: 249
i: [1,125,0,9,3,0,0,1] d19_relat_1
s: 125
d: [2,253,1,6,2,0,40,125,135]
s: 253
i: [1,133,0,9,3,0,0,255] t5_subset
s: 133
i: [1,132,0,10,3,0,0,255] t4_subset
s: 132
d: [2,246,2,6,2,0,40,244,136]
s: 246
d: [2,243,1,7,3,0,40,109,129]
s: 243
d: [2,257,1,7,2,0,40,133,134]
s: 257
d: [2,258,1,8,3,0,40,133,129]
s: 258
d: [2,259,1,8,2,0,40,132,134]
s: 259
d: [2,268,2,5,2,0,40,258,107]
a: [3,276,2,2,1,268]
s: 276
r: 201
r: 191
r: 176
r: 117
r: 118
r: 258
r: 257
r: 236
r: 235
r: 234
r: 233
r: 232
r: 230
r: 133
r: 131
r: 130
r: 114
r: 113
r: 112
r: 110
r: 108
d: [2,284,3,0,0,2,40,276,181]
e: 284
d: [2,283,3,0,0,2,40,276,161]
e: 283
d: [2,282,3,0,0,2,52,181,276]
e: 282
r: 276
s: 176
s: 232
s: 191
s: 201
a: [3,279,2,3,1,268]
s: 279
s: 108
s: 110
s: 112
s: 113
s: 114
s: 131
s: 230
s: 117
s: 118
s: 233
s: 130
s: 234
s: 133
s: 235
s: 236
d: [2,294,1,5,2,0,44,107,112]
s: 294
d: [2,300,1,5,2,0,40,113,110]
s: 300
d: [2,301,1,5,2,0,40,113,108]
s: 301
d: [2,302,1,5,2,0,40,114,110]
s: 302
d: [2,303,1,5,2,0,40,114,108]
s: 303
s: 257
d: [2,299,3,5,2,1,44,279,112]
s: 299
d: [2,293,1,7,3,0,44,112,112]
s: 293
d: [2,310,1,7,2,0,40,133,134]
s: 310
d: [2,267,2,6,2,0,40,257,244]
s: 267
s: 258
d: [2,312,2,6,2,0,40,257,244]
s: 312
d: [2,311,1,8,3,0,40,133,129]
s: 311
d: [2,260,1,9,3,0,40,132,129]
s: 260
d: [2,313,3,6,2,0,40,267,110]
s: 313
d: [2,245,1,10,3,0,40,127,129]
s: 245
d: [2,314,3,6,2,0,40,267,108]
s: 314
d: [2,251,2,8,3,0,40,123,230]
s: 251
d: [2,255,2,8,3,0,40,125,230]
s: 255
d: [2,330,3,7,3,0,40,251,114]
d: [2,331,3,5,2,0,21,330,110]
s: 331
d: [2,334,4,0,0,2,52,161,147,331]
e: 334
d: [2,335,4,2,1,1,40,331,147]
d: [2,339,4,0,0,2,21,335,161]
e: 339
r: 232
r: 147
a: [3,146,0,3,1,103]
s: 146
a: [3,151,0,3,1,103]
s: 151
d: [2,342,1,5,2,1,44,151,112]
s: 342
d: [2,305,2,8,3,0,40,230,123]
s: 305
d: [2,332,3,7,3,0,40,255,113]
d: [2,333,3,5,2,0,21,332,110]
s: 333
d: [2,343,4,0,0,2,52,161,151,333]
e: 343
d: [2,344,4,2,1,1,40,333,151]
d: [2,348,4,0,0,2,21,344,161]
e: 348
f: 348,339,105,103
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__l186_relat_1
% SZS output start Proof for relat_1__l186_relat_1
1. ! [X0] : ! [X1] : (v5_relat_1(k1_xboole_0,X1) & v4_relat_1(k1_xboole_0,X0)) [input l186_relat_1]
2. ~! [X0] : ! [X1] : (v5_relat_1(k1_xboole_0,X1) & v4_relat_1(k1_xboole_0,X0)) [negated conjecture 1]
4. ! [X0] : (v1_xboole_0(X0) => v1_relat_1(X0)) [input cc1_relat_1]
9. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0))) [input d18_relat_1]
10. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v5_relat_1(X1,X0) <=> r1_tarski(k10_xtuple_0(X1),X0))) [input d19_relat_1]
19. ! [X0] : (v1_xboole_0(X0) => v1_xboole_0(k9_xtuple_0(X0))) [input fc10_relat_1]
20. ! [X0] : (v1_xboole_0(X0) => v1_xboole_0(k10_xtuple_0(X0))) [input fc11_relat_1]
21. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
36. ! [X0] : r1_tarski(k1_xboole_0,X0) [input t2_xboole_1]
40. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
43. ~! [X0,X1] : (v5_relat_1(k1_xboole_0,X1) & v4_relat_1(k1_xboole_0,X0)) [flattening 2]
44. ! [X0,X1] : (v1_relat_1(X1) => (v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0))) [flattening 9]
45. ! [X0,X1] : (v1_relat_1(X1) => (v5_relat_1(X1,X0) <=> r1_tarski(k10_xtuple_0(X1),X0))) [flattening 10]
55. ? [X0,X1] : (~v5_relat_1(k1_xboole_0,X1) | ~v4_relat_1(k1_xboole_0,X0)) [ennf transformation 43]
58. ! [X0] : (v1_relat_1(X0) | ~v1_xboole_0(X0)) [ennf transformation 4]
60. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 40]
61. ! [X0] : (v1_xboole_0(k10_xtuple_0(X0)) | ~v1_xboole_0(X0)) [ennf transformation 20]
62. ! [X0] : (v1_xboole_0(k9_xtuple_0(X0)) | ~v1_xboole_0(X0)) [ennf transformation 19]
75. ! [X0,X1] : ((v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0)) | ~v1_relat_1(X1)) [ennf transformation 44]
76. ! [X0,X1] : ((v5_relat_1(X1,X0) <=> r1_tarski(k10_xtuple_0(X1),X0)) | ~v1_relat_1(X1)) [ennf transformation 45]
86. ? [X0,X1] : (~v5_relat_1(k1_xboole_0,X1) | ~v4_relat_1(k1_xboole_0,X0)) => (~v5_relat_1(k1_xboole_0,sK1) | ~v4_relat_1(k1_xboole_0,sK0)) [choice axiom]
87. ~v5_relat_1(k1_xboole_0,sK1) | ~v4_relat_1(k1_xboole_0,sK0) [skolemisation 55,86]
88. ! [X0,X1] : (((v4_relat_1(X1,X0) | ~r1_tarski(k9_xtuple_0(X1),X0)) & (r1_tarski(k9_xtuple_0(X1),X0) | ~v4_relat_1(X1,X0))) | ~v1_relat_1(X1)) [nnf transformation 75]
89. ! [X0,X1] : (((v5_relat_1(X1,X0) | ~r1_tarski(k10_xtuple_0(X1),X0)) & (r1_tarski(k10_xtuple_0(X1),X0) | ~v5_relat_1(X1,X0))) | ~v1_relat_1(X1)) [nnf transformation 76]
103. ~v5_relat_1(k1_xboole_0,sK1) | ~v4_relat_1(k1_xboole_0,sK0) [cnf transformation 87] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
105. v1_xboole_0(k1_xboole_0) [cnf transformation 21] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
107. r1_tarski(k1_xboole_0,X0) [cnf transformation 36] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
110. ~v1_xboole_0(X0) | v1_relat_1(X0) [cnf transformation 58] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
112. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 60] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
113. v1_xboole_0(k10_xtuple_0(X0)) | ~v1_xboole_0(X0) [cnf transformation 61] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
114. v1_xboole_0(k9_xtuple_0(X0)) | ~v1_xboole_0(X0) [cnf transformation 62] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
123. ~r1_tarski(k9_xtuple_0(X1),X0) | v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 88] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
125. ~r1_tarski(k10_xtuple_0(X1),X0) | v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 89] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
145. 1 <=> v4_relat_1(k1_xboole_0,sK0) [avatar definition]
147. ~v4_relat_1(k1_xboole_0,sK0) <- (~1) [avatar component clause 145] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
149. 2 <=> v5_relat_1(k1_xboole_0,sK1) [avatar definition]
151. ~v5_relat_1(k1_xboole_0,sK1) <- (~2) [avatar component clause 149] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
152. ~1 | ~2 [avatar split clause 103,149,145]
159. 4 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
161. v1_xboole_0(k1_xboole_0) <- (4) [avatar component clause 159] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
162. 4 [avatar split clause 105,159]
230. r1_tarski(X0,X1) | ~v1_xboole_0(X0) [superposition 107,112] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
251. ~v1_xboole_0(k9_xtuple_0(X3)) | ~v1_relat_1(X3) | v4_relat_1(X3,X4) [resolution 123,230] {a:2,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
255. ~v1_xboole_0(k10_xtuple_0(X3)) | ~v1_relat_1(X3) | v5_relat_1(X3,X4) [resolution 125,230] {a:2,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
330. ~v1_relat_1(X0) | v4_relat_1(X0,X1) | ~v1_xboole_0(X0) [resolution 251,114] {a:3,w:7,wCS:2800,sine:1,thAx:0,allAx:4,thDist:-4}
331. v4_relat_1(X0,X1) | ~v1_xboole_0(X0) [subsumption resolution 330,110] {a:3,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
332. ~v1_relat_1(X0) | v5_relat_1(X0,X1) | ~v1_xboole_0(X0) [resolution 255,113] {a:3,w:7,wCS:2800,sine:1,thAx:0,allAx:4,thDist:-4}
333. v5_relat_1(X0,X1) | ~v1_xboole_0(X0) [subsumption resolution 332,110] {a:3,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
335. ~v1_xboole_0(k1_xboole_0) <- (~1) [resolution 331,147] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
339. $false <- (~1, 4) [subsumption resolution 335,161] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
340. 1 | ~4 [avatar contradiction clause 339]
344. ~v1_xboole_0(k1_xboole_0) <- (~2) [resolution 333,151] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
348. $false <- (~2, 4) [subsumption resolution 344,161] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
349. 2 | ~4 [avatar contradiction clause 348]
350. $false [avatar sat refutation 152,162,340,349] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:94,thDist:-94}
% SZS output end Proof for relat_1__l186_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 41
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 194
% Active clauses: 92
% Passive clauses: 118
% Final active clauses: 67
% Final passive clauses: 11
% Main loop iterations started: 107
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 44
% Backward subsumptions: 21
% 
% >>> Generating Inferences
% Binary resolution: 86
% Unit resulting resolution: 3
% Forward superposition: 3
% Backward superposition: 10
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 25
% SAT solver unit clauses: 10
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

