i: [1,52,0,2,1,1,0,0] t75_funct_4
a: [3,82,0,2,1,52]
s: 82
i: [1,51,0,2,1,1,0,0] t75_funct_4
a: [3,87,0,2,1,51]
s: 87
i: [1,53,0,7,1,1,0,0] t75_funct_4
a: [3,77,0,7,1,53]
s: 77
i: [1,73,0,2,1,0,0,1] rc1_funct_1
a: [3,100,0,2,1,73]
s: 100
i: [1,72,0,2,1,0,0,1] rc1_funct_1
a: [3,105,0,2,1,72]
s: 105
i: [1,61,0,3,1,0,0,255] reflexivity_r1_tarski
s: 61
i: [1,60,0,4,1,0,0,255] existence_m1_subset_1
s: 60
i: [1,62,0,6,2,0,0,1] dt_k5_relat_1
s: 62
i: [1,57,0,7,2,0,0,255] t59_relat_1
s: 57
i: [1,58,0,7,2,0,0,255] t3_subset
s: 58
i: [1,59,0,7,2,0,0,255] t3_subset
s: 59
d: [2,107,1,5,1,0,40,58,60]
s: 107
i: [1,54,0,8,3,0,0,255] cc2_relat_1
s: 54
d: [2,109,1,6,2,0,40,54,60]
s: 109
i: [1,70,0,8,3,0,0,1] fc8_funct_1
s: 70
i: [1,71,0,8,3,0,0,1] fc8_funct_1
s: 71
d: [2,110,1,7,3,0,40,54,59]
s: 110
i: [1,64,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,94,0,9,3,64]
s: 94
i: [1,55,0,10,4,0,0,255] cc3_funct_1
s: 55
i: [1,63,0,11,2,0,0,1] rd5_relat_1
s: 63
i: [1,66,0,12,5,0,0,1] dt_k1_funct_4
s: 66
d: [2,115,1,8,3,0,40,55,60]
s: 115
i: [1,67,0,12,5,0,0,1] dt_k1_funct_4
s: 67
d: [2,116,1,9,4,0,40,55,59]
s: 116
i: [1,68,0,12,5,0,0,1] dt_k1_funct_4
s: 68
i: [1,69,0,12,5,0,0,1] dt_k1_funct_4
s: 69
i: [1,56,0,16,6,0,0,255] l80_funct_4
s: 56
d: [2,129,0,14,5,0,21,56,116]
d: [2,130,0,12,4,0,21,129,110]
s: 130
d: [2,131,1,12,4,1,44,77,130]
d: [2,136,1,9,3,1,19,131]
d: [2,137,1,7,2,2,21,136,82]
d: [2,138,1,5,1,3,21,137,87]
a: [3,142,1,5,1,138]
s: 142
d: [2,144,2,0,0,2,52,87,142,57]
e: 144
d: [2,145,2,2,1,1,40,142,57]
d: [2,147,2,0,0,2,21,145,87]
e: 147
f: 147,138,51,52,53
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t75_funct_4
% SZS output start Proof for funct_4__t75_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : k1_funct_4(X0,k5_relat_1(X0,X1)) = X0) [input t75_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : k1_funct_4(X0,k5_relat_1(X0,X1)) = X0) [negated conjecture 1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
7. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_funct_1(X1))) [input cc3_funct_1]
17. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
18. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) => k1_funct_4(X1,X0) = X1))) [input l80_funct_4]
19. ! [X0] : ! [X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [input t59_relat_1]
20. ! [X0,X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [flattening 19]
21. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 17]
23. ? [X0] : (? [X1] : k1_funct_4(X0,k5_relat_1(X0,X1)) != X0 & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
24. ? [X0] : (? [X1] : k1_funct_4(X0,k5_relat_1(X0,X1)) != X0 & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 23]
25. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 6]
26. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 7]
27. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 26]
28. ! [X0] : (! [X1] : ((k1_funct_4(X1,X0) = X1 | ~r1_tarski(X0,X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 18]
29. ! [X0] : (! [X1] : (k1_funct_4(X1,X0) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 28]
30. ! [X0,X1] : (r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1)) [ennf transformation 20]
43. ? [X0] : (? [X1] : k1_funct_4(X0,k5_relat_1(X0,X1)) != X0 & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : sK0 != k1_funct_4(sK0,k5_relat_1(sK0,X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
44. ? [X1] : sK0 != k1_funct_4(sK0,k5_relat_1(sK0,X1)) => sK0 != k1_funct_4(sK0,k5_relat_1(sK0,sK1)) [choice axiom]
45. sK0 != k1_funct_4(sK0,k5_relat_1(sK0,sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 24,44,43]
46. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 21]
51. v1_relat_1(sK0) [cnf transformation 45] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. v1_funct_1(sK0) [cnf transformation 45] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
53. sK0 != k1_funct_4(sK0,k5_relat_1(sK0,sK1)) [cnf transformation 45] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 25] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
55. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_funct_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
56. k1_funct_4(X1,X0) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 29] {a:0,w:16,wCS:6400,sine:255,thAx:0,allAx:1,thDist:-1}
57. r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1) [cnf transformation 30] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
59. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 46] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
75. 1 <=> sK0 = k1_funct_4(sK0,k5_relat_1(sK0,sK1)) [avatar definition]
77. sK0 != k1_funct_4(sK0,k5_relat_1(sK0,sK1)) <- (~1) [avatar component clause 75] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
78. ~1 [avatar split clause 53,75]
80. 2 <=> v1_funct_1(sK0) [avatar definition]
82. v1_funct_1(sK0) <- (2) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
83. 2 [avatar split clause 52,80]
85. 3 <=> v1_relat_1(sK0) [avatar definition]
87. v1_relat_1(sK0) <- (3) [avatar component clause 85] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
88. 3 [avatar split clause 51,85]
110. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 54,59] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. ~r1_tarski(X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | v1_funct_1(X1) [resolution 55,59] {a:1,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
129. k1_funct_4(X1,X0) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [subsumption resolution 56,116] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
130. k1_funct_4(X1,X0) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [subsumption resolution 129,110] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
131. sK0 != sK0 | ~r1_tarski(k5_relat_1(sK0,sK1),sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [superposition 77,130] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
136. ~r1_tarski(k5_relat_1(sK0,sK1),sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [trivial inequality removal 131] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
137. ~r1_tarski(k5_relat_1(sK0,sK1),sK0) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 136,82] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
138. ~r1_tarski(k5_relat_1(sK0,sK1),sK0) <- (~1, 2, 3) [subsumption resolution 137,87] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
140. 8 <=> r1_tarski(k5_relat_1(sK0,sK1),sK0) [avatar definition]
142. ~r1_tarski(k5_relat_1(sK0,sK1),sK0) <- (~8) [avatar component clause 140] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
143. ~8 | 1 | ~2 | ~3 [avatar split clause 138,85,80,75,140]
145. ~v1_relat_1(sK0) <- (~8) [resolution 142,57] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
147. $false <- (3, ~8) [subsumption resolution 145,87] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
148. ~3 | 8 [avatar contradiction clause 147]
149. $false [avatar sat refutation 78,83,88,143,148] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:45,thDist:-45}
% SZS output end Proof for funct_4__t75_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 65
% Active clauses: 25
% Passive clauses: 29
% Final active clauses: 25
% Main loop iterations started: 29
% 
% >>> Simplifying Inferences
% Duplicate literals: 14
% Trivial inequalities: 1
% Fw subsumption resolutions: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 6
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 5
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

