i: [1,53,0,8,2,1,0,0] t90_zfmisc_1
a: [3,88,0,3,1,53]
s: 88
a: [3,85,0,5,1,53]
s: 85
d: [2,122,0,5,1,2,22,85,88]
a: [3,126,0,5,1,122]
s: 126
i: [1,54,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,100,0,2,1,54]
s: 100
i: [1,55,0,2,1,0,0,1] fc1_xboole_0
a: [3,105,0,2,1,55]
s: 105
i: [1,75,0,2,1,0,0,2] rc2_xboole_0
a: [3,115,0,2,1,75]
s: 115
i: [1,76,0,2,1,0,0,2] rc1_xboole_0
a: [3,120,0,2,1,76]
s: 120
i: [1,57,0,3,1,0,0,255] fc2_xboole_0
s: 57
i: [1,56,0,3,1,0,0,255] d2_xboole_0
a: [3,110,0,3,1,56]
s: 110
r: 100
i: [1,71,0,4,1,0,0,2] fc1_zfmisc_1
s: 71
i: [1,72,0,4,1,0,0,255] fc3_xboole_0
s: 72
i: [1,58,0,5,2,0,0,1] t6_boole
s: 58
i: [1,62,0,5,2,0,0,2] t7_boole
s: 62
d: [2,131,1,7,2,1,44,126,58]
s: 131
i: [1,74,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 74
i: [1,59,0,7,2,0,0,2] t7_xboole_0
s: 59
i: [1,61,0,7,3,0,0,1] t8_boole
s: 61
i: [1,73,0,7,1,0,0,255] commutativity_k2_tarski
s: 73
d: [2,129,1,7,3,0,44,58,58]
s: 129
i: [1,60,0,10,1,0,0,255] d5_tarski
s: 60
i: [1,80,0,11,2,0,0,1] d2_zfmisc_1
s: 80
d: [2,132,1,7,2,0,40,59,74]
s: 132
i: [1,81,0,11,2,0,0,1] d2_zfmisc_1
s: 81
d: [2,145,1,7,2,0,40,80,62]
s: 145
i: [1,78,0,13,3,0,0,1] d2_zfmisc_1
s: 78
i: [1,79,0,16,2,0,0,1] d2_zfmisc_1
s: 79
d: [2,147,1,7,2,0,40,81,62]
s: 147
i: [1,67,0,17,3,0,0,1] d2_zfmisc_1
s: 67
d: [2,155,1,7,2,0,44,71,79]
s: 155
i: [1,68,0,17,3,0,0,1] d2_zfmisc_1
s: 68
i: [1,69,0,25,3,0,0,1] d2_zfmisc_1
s: 69
d: [2,148,2,7,2,0,40,145,59]
s: 148
d: [2,176,3,0,0,2,52,105,126,148]
e: 176
d: [2,177,3,5,2,1,44,126,148]
d: [2,185,3,2,1,1,19,177]
d: [2,187,3,0,0,2,21,185,105]
e: 187
r: 88
r: 131
r: 126
a: [3,89,0,3,1,53]
s: 89
i: [1,52,0,8,2,1,0,0] t90_zfmisc_1
a: [3,93,0,3,1,52]
s: 93
s: 85
a: [3,125,0,5,1,122]
s: 125
r: 93
a: [3,94,0,3,1,52]
s: 94
a: [3,84,0,5,1,53]
s: 84
d: [2,215,2,5,2,1,44,147,84]
a: [3,222,2,2,1,215]
s: 222
s: 125
d: [2,202,0,5,1,2,22,85,93]
a: [3,206,0,5,1,202]
s: 206
d: [2,217,2,5,2,1,44,145,84]
a: [3,227,2,2,1,217]
s: 227
i: [1,70,0,25,5,0,0,1] d2_zfmisc_1
s: 70
d: [2,229,1,7,2,1,44,125,58]
s: 229
d: [2,235,1,7,2,1,44,206,58]
s: 235
d: [2,216,1,11,3,1,44,78,84]
s: 216
d: [2,256,2,8,3,1,40,216,62]
d: [2,260,2,6,2,2,21,256,105]
a: [3,263,2,3,1,260]
s: 263
r: 216
d: [2,270,3,0,0,2,52,94,263,59]
e: 270
d: [2,271,3,3,1,1,40,263,59]
d: [2,277,3,0,0,2,21,271,94]
e: 277
r: 263
s: 216
a: [3,266,2,3,1,260]
s: 266
r: 216
d: [2,288,3,0,0,2,52,89,266,59]
e: 288
d: [2,289,3,3,1,1,40,266,59]
d: [2,295,3,0,0,2,21,289,89]
e: 295
r: 84
r: 94
r: 222
r: 266
s: 93
r: 229
r: 125
s: 85
a: [3,221,2,2,1,215]
s: 221
d: [2,208,0,5,1,2,22,125,93]
a: [3,212,0,5,1,208]
s: 212
d: [2,299,1,7,2,2,23,229,93]
s: 299
d: [2,304,1,7,2,1,44,212,58]
s: 304
d: [2,232,1,11,3,1,44,78,125]
s: 232
d: [2,182,3,7,3,0,44,145,148]
d: [2,183,3,5,2,0,18,182]
a: [3,192,3,2,1,183]
s: 192
r: 115
r: 227
r: 57
r: 71
r: 72
r: 304
r: 299
r: 235
r: 155
r: 148
r: 147
r: 145
r: 62
r: 61
r: 58
d: [2,338,4,0,0,2,40,192,105]
e: 338
d: [2,337,4,0,0,2,40,192,120]
e: 337
d: [2,336,4,0,0,2,52,120,192]
e: 336
r: 192
s: 115
s: 227
s: 57
s: 71
s: 235
s: 72
s: 299
s: 58
s: 62
s: 304
s: 61
d: [2,352,1,7,2,1,44,212,58]
s: 352
d: [2,307,1,11,3,1,44,78,212]
s: 307
d: [2,335,1,11,3,2,22,232,93]
s: 335
d: [2,310,2,9,3,2,44,299,58]
s: 310
s: 145
d: [2,351,2,9,3,2,44,299,58]
s: 351
s: 147
s: 155
d: [2,366,2,11,3,1,40,307,74]
s: 366
d: [2,342,1,7,2,0,44,71,79]
s: 342
a: [3,195,3,3,1,183]
s: 195
r: 307
r: 366
d: [2,413,4,0,0,2,52,206,195,195,67]
e: 413
d: [2,414,4,0,0,2,52,206,195,195,67]
e: 414
i: [1,51,0,11,3,1,0,0] t90_zfmisc_1
d: [2,180,3,7,3,0,44,147,148]
d: [2,283,2,13,4,1,44,216,58]
d: [2,286,2,8,3,1,21,283,62]
f: 414,336,295,286,277,202,180,187,122,76,55,51,52,53
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t90_zfmisc_1
% SZS output start Proof for zfmisc_1__t90_zfmisc_1
1. ! [X0] : ! [X1] : (k2_zfmisc_1(X0,X1) = k1_xboole_0 <=> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [input t90_zfmisc_1]
2. ~! [X0] : ! [X1] : (k2_zfmisc_1(X0,X1) = k1_xboole_0 <=> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [negated conjecture 1]
6. ! [X0] : ! [X1] : ! [X2] : (k2_zfmisc_1(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> ? [X4] : ? [X5] : (k4_tarski(X4,X5) = X3 & r2_hidden(X5,X1) & r2_hidden(X4,X0)))) [input d2_zfmisc_1]
14. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
18. ? [X0] : v1_xboole_0(X0) [input rc1_xboole_0]
20. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
21. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
22. ! [X0] : ~(! [X1] : ~r2_hidden(X1,X0) & k1_xboole_0 != X0) [input t7_xboole_0]
24. ~! [X0,X1] : (k2_zfmisc_1(X0,X1) = k1_xboole_0 <=> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [flattening 2]
27. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 21]
28. ! [X0,X1,X2] : (k2_zfmisc_1(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> ? [X4,X5] : (k4_tarski(X4,X5) = X3 & r2_hidden(X5,X1) & r2_hidden(X4,X0)))) [flattening 6]
29. ? [X0,X1] : (k2_zfmisc_1(X0,X1) = k1_xboole_0 <~> (k1_xboole_0 = X1 | k1_xboole_0 = X0)) [ennf transformation 24]
30. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 20]
31. ! [X0] : (? [X1] : r2_hidden(X1,X0) | k1_xboole_0 = X0) [ennf transformation 22]
33. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 27]
35. ? [X0,X1] : (((k1_xboole_0 != X1 & k1_xboole_0 != X0) | k2_zfmisc_1(X0,X1) != k1_xboole_0) & ((k1_xboole_0 = X1 | k1_xboole_0 = X0) | k2_zfmisc_1(X0,X1) = k1_xboole_0)) [nnf transformation 29]
36. ? [X0,X1] : (((k1_xboole_0 != X1 & k1_xboole_0 != X0) | k2_zfmisc_1(X0,X1) != k1_xboole_0) & (k1_xboole_0 = X1 | k1_xboole_0 = X0 | k2_zfmisc_1(X0,X1) = k1_xboole_0)) [flattening 35]
37. ? [X0,X1] : (((k1_xboole_0 != X1 & k1_xboole_0 != X0) | k2_zfmisc_1(X0,X1) != k1_xboole_0) & (k1_xboole_0 = X1 | k1_xboole_0 = X0 | k2_zfmisc_1(X0,X1) = k1_xboole_0)) => (((k1_xboole_0 != sK1 & k1_xboole_0 != sK0) | k1_xboole_0 != k2_zfmisc_1(sK0,sK1)) & (k1_xboole_0 = sK1 | k1_xboole_0 = sK0 | k1_xboole_0 = k2_zfmisc_1(sK0,sK1))) [choice axiom]
38. ((k1_xboole_0 != sK1 & k1_xboole_0 != sK0) | k1_xboole_0 != k2_zfmisc_1(sK0,sK1)) & (k1_xboole_0 = sK1 | k1_xboole_0 = sK0 | k1_xboole_0 = k2_zfmisc_1(sK0,sK1)) [skolemisation 36,37]
39. ! [X0] : (? [X1] : r2_hidden(X1,X0) => r2_hidden(sK2(X0),X0)) [choice axiom]
40. ! [X0] : (r2_hidden(sK2(X0),X0) | k1_xboole_0 = X0) [skolemisation 31,39]
41. ! [X0,X1,X2] : ((k2_zfmisc_1(X0,X1) = X2 | ? [X3] : ((! [X4,X5] : (k4_tarski(X4,X5) != X3 | ~r2_hidden(X5,X1) | ~r2_hidden(X4,X0)) | ~r2_hidden(X3,X2)) & (? [X4,X5] : (k4_tarski(X4,X5) = X3 & r2_hidden(X5,X1) & r2_hidden(X4,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | ! [X4,X5] : (k4_tarski(X4,X5) != X3 | ~r2_hidden(X5,X1) | ~r2_hidden(X4,X0))) & (? [X4,X5] : (k4_tarski(X4,X5) = X3 & r2_hidden(X5,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X3,X2))) | k2_zfmisc_1(X0,X1) != X2)) [nnf transformation 28]
42. ! [X0,X1,X2] : ((k2_zfmisc_1(X0,X1) = X2 | ? [X3] : ((! [X4,X5] : (k4_tarski(X4,X5) != X3 | ~r2_hidden(X5,X1) | ~r2_hidden(X4,X0)) | ~r2_hidden(X3,X2)) & (? [X6,X7] : (k4_tarski(X6,X7) = X3 & r2_hidden(X7,X1) & r2_hidden(X6,X0)) | r2_hidden(X3,X2)))) & (! [X8] : ((r2_hidden(X8,X2) | ! [X9,X10] : (k4_tarski(X9,X10) != X8 | ~r2_hidden(X10,X1) | ~r2_hidden(X9,X0))) & (? [X11,X12] : (k4_tarski(X11,X12) = X8 & r2_hidden(X12,X1) & r2_hidden(X11,X0)) | ~r2_hidden(X8,X2))) | k2_zfmisc_1(X0,X1) != X2)) [rectify 41]
43. ! [X2,X1,X0] : (? [X3] : ((! [X4,X5] : (k4_tarski(X4,X5) != X3 | ~r2_hidden(X5,X1) | ~r2_hidden(X4,X0)) | ~r2_hidden(X3,X2)) & (? [X6,X7] : (k4_tarski(X6,X7) = X3 & r2_hidden(X7,X1) & r2_hidden(X6,X0)) | r2_hidden(X3,X2))) => ((! [X5,X4] : (k4_tarski(X4,X5) != sK3(X0,X1,X2) | ~r2_hidden(X5,X1) | ~r2_hidden(X4,X0)) | ~r2_hidden(sK3(X0,X1,X2),X2)) & (? [X7,X6] : (k4_tarski(X6,X7) = sK3(X0,X1,X2) & r2_hidden(X7,X1) & r2_hidden(X6,X0)) | r2_hidden(sK3(X0,X1,X2),X2)))) [choice axiom]
44. ! [X2,X1,X0] : (? [X7,X6] : (k4_tarski(X6,X7) = sK3(X0,X1,X2) & r2_hidden(X7,X1) & r2_hidden(X6,X0)) => (sK3(X0,X1,X2) = k4_tarski(sK4(X0,X1,X2),sK5(X0,X1,X2)) & r2_hidden(sK5(X0,X1,X2),X1) & r2_hidden(sK4(X0,X1,X2),X0))) [choice axiom]
45. ! [X8,X1,X0] : (? [X11,X12] : (k4_tarski(X11,X12) = X8 & r2_hidden(X12,X1) & r2_hidden(X11,X0)) => (k4_tarski(sK6(X0,X1,X8),sK7(X0,X1,X8)) = X8 & r2_hidden(sK7(X0,X1,X8),X1) & r2_hidden(sK6(X0,X1,X8),X0))) [choice axiom]
46. ! [X0,X1,X2] : ((k2_zfmisc_1(X0,X1) = X2 | ((! [X4,X5] : (k4_tarski(X4,X5) != sK3(X0,X1,X2) | ~r2_hidden(X5,X1) | ~r2_hidden(X4,X0)) | ~r2_hidden(sK3(X0,X1,X2),X2)) & ((sK3(X0,X1,X2) = k4_tarski(sK4(X0,X1,X2),sK5(X0,X1,X2)) & r2_hidden(sK5(X0,X1,X2),X1) & r2_hidden(sK4(X0,X1,X2),X0)) | r2_hidden(sK3(X0,X1,X2),X2)))) & (! [X8] : ((r2_hidden(X8,X2) | ! [X9,X10] : (k4_tarski(X9,X10) != X8 | ~r2_hidden(X10,X1) | ~r2_hidden(X9,X0))) & ((k4_tarski(sK6(X0,X1,X8),sK7(X0,X1,X8)) = X8 & r2_hidden(sK7(X0,X1,X8),X1) & r2_hidden(sK6(X0,X1,X8),X0)) | ~r2_hidden(X8,X2))) | k2_zfmisc_1(X0,X1) != X2)) [skolemisation 42,45,44,43]
49. ? [X0] : v1_xboole_0(X0) => v1_xboole_0(sK9) [choice axiom]
50. v1_xboole_0(sK9) [skolemisation 18,49]
51. k1_xboole_0 = sK1 | k1_xboole_0 = sK0 | k1_xboole_0 = k2_zfmisc_1(sK0,sK1) [cnf transformation 38] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. k1_xboole_0 != sK0 | k1_xboole_0 != k2_zfmisc_1(sK0,sK1) [cnf transformation 38] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. k1_xboole_0 != sK1 | k1_xboole_0 != k2_zfmisc_1(sK0,sK1) [cnf transformation 38] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. v1_xboole_0(k1_xboole_0) [cnf transformation 14] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
58. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 30] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. r2_hidden(sK2(X0),X0) | k1_xboole_0 = X0 [cnf transformation 40] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
62. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 33] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
63. r2_hidden(sK6(X0,X1,X8),X0) | ~r2_hidden(X8,X2) | k2_zfmisc_1(X0,X1) != X2 [cnf transformation 46] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:4,thDist:-4}
64. r2_hidden(sK7(X0,X1,X8),X1) | ~r2_hidden(X8,X2) | k2_zfmisc_1(X0,X1) != X2 [cnf transformation 46] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:4,thDist:-4}
66. r2_hidden(X8,X2) | k4_tarski(X9,X10) != X8 | ~r2_hidden(X10,X1) | ~r2_hidden(X9,X0) | k2_zfmisc_1(X0,X1) != X2 [cnf transformation 46] {a:0,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
67. r2_hidden(sK4(X0,X1,X2),X0) | k2_zfmisc_1(X0,X1) = X2 | r2_hidden(sK3(X0,X1,X2),X2) [cnf transformation 46] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
76. v1_xboole_0(sK9) [cnf transformation 50] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
77. r2_hidden(k4_tarski(X9,X10),X2) | ~r2_hidden(X10,X1) | ~r2_hidden(X9,X0) | k2_zfmisc_1(X0,X1) != X2 [equality resolution 66] {a:0,w:16,wCS:6400,sine:1,thAx:0,allAx:4,thDist:-4}
78. r2_hidden(k4_tarski(X9,X10),k2_zfmisc_1(X0,X1)) | ~r2_hidden(X10,X1) | ~r2_hidden(X9,X0) [equality resolution 77] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
80. r2_hidden(sK7(X0,X1,X8),X1) | ~r2_hidden(X8,k2_zfmisc_1(X0,X1)) [equality resolution 64] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
81. r2_hidden(sK6(X0,X1,X8),X0) | ~r2_hidden(X8,k2_zfmisc_1(X0,X1)) [equality resolution 63] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
83. 1 <=> k1_xboole_0 = k2_zfmisc_1(sK0,sK1) [avatar definition]
84. k1_xboole_0 = k2_zfmisc_1(sK0,sK1) <- (1) [avatar component clause 83] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
85. k1_xboole_0 != k2_zfmisc_1(sK0,sK1) <- (~1) [avatar component clause 83] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. 2 <=> k1_xboole_0 = sK1 [avatar definition]
88. k1_xboole_0 = sK1 <- (2) [avatar component clause 87] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
89. k1_xboole_0 != sK1 <- (~2) [avatar component clause 87] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
90. ~1 | ~2 [avatar split clause 53,87,83]
92. 3 <=> k1_xboole_0 = sK0 [avatar definition]
93. k1_xboole_0 = sK0 <- (3) [avatar component clause 92] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
94. k1_xboole_0 != sK0 <- (~3) [avatar component clause 92] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
95. ~1 | ~3 [avatar split clause 52,92,83]
96. 1 | 3 | 2 [avatar split clause 51,87,92,83]
103. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
105. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 103] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
106. 5 [avatar split clause 55,103]
118. 8 <=> v1_xboole_0(sK9) [avatar definition]
120. v1_xboole_0(sK9) <- (8) [avatar component clause 118] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
121. 8 [avatar split clause 76,118]
122. k1_xboole_0 != k2_zfmisc_1(sK0,k1_xboole_0) <- (~1, 2) [forward demodulation 85,88] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
124. 9 <=> k1_xboole_0 = k2_zfmisc_1(sK0,k1_xboole_0) [avatar definition]
126. k1_xboole_0 != k2_zfmisc_1(sK0,k1_xboole_0) <- (~9) [avatar component clause 124] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
127. ~9 | 1 | ~2 [avatar split clause 122,87,83,124]
145. ~r2_hidden(X3,k2_zfmisc_1(X4,X5)) | ~v1_xboole_0(X5) [resolution 80,62] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
147. ~r2_hidden(X3,k2_zfmisc_1(X4,X5)) | ~v1_xboole_0(X4) [resolution 81,62] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
148. k1_xboole_0 = k2_zfmisc_1(X1,X0) | ~v1_xboole_0(X0) [resolution 145,59] {a:2,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:7,thDist:-7}
177. k1_xboole_0 != k1_xboole_0 | ~v1_xboole_0(k1_xboole_0) <- (~9) [superposition 126,148] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
180. ~r2_hidden(X6,k1_xboole_0) | ~v1_xboole_0(X4) | ~v1_xboole_0(X5) [superposition 147,148] {a:3,w:7,wCS:2800,sine:1,thAx:0,allAx:12,thDist:-12}
185. ~v1_xboole_0(k1_xboole_0) <- (~9) [trivial inequality removal 177] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
187. $false <- (5, ~9) [subsumption resolution 185,105] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
188. ~5 | 9 [avatar contradiction clause 187]
191. 10 <=> ! [X12] : ~v1_xboole_0(X12) [avatar definition]
192. ~v1_xboole_0(X12) <- (10) [avatar component clause 191] {a:3,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:12,thDist:-12}
194. 11 <=> ! [X13] : ~r2_hidden(X13,k1_xboole_0) [avatar definition]
195. ~r2_hidden(X13,k1_xboole_0) <- (11) [avatar component clause 194] {a:3,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:12,thDist:-12}
197. 10 | 10 | 11 [avatar split clause 180,194,191,191]
202. k1_xboole_0 != k2_zfmisc_1(k1_xboole_0,sK1) <- (~1, 3) [forward demodulation 85,93] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
204. 13 <=> k1_xboole_0 = k2_zfmisc_1(k1_xboole_0,sK1) [avatar definition]
206. k1_xboole_0 != k2_zfmisc_1(k1_xboole_0,sK1) <- (~13) [avatar component clause 204] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
207. ~13 | 1 | ~3 [avatar split clause 202,92,83,204]
216. r2_hidden(k4_tarski(X2,X3),k1_xboole_0) | ~r2_hidden(X3,sK1) | ~r2_hidden(X2,sK0) <- (1) [superposition 78,84] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
262. 17 <=> ! [X3] : ~r2_hidden(X3,sK0) [avatar definition]
263. ~r2_hidden(X3,sK0) <- (17) [avatar component clause 262] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
265. 18 <=> ! [X2] : ~r2_hidden(X2,sK1) [avatar definition]
266. ~r2_hidden(X2,sK1) <- (18) [avatar component clause 265] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
271. k1_xboole_0 = sK0 <- (17) [resolution 263,59] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
277. $false <- (~3, 17) [subsumption resolution 271,94] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
278. 3 | ~17 [avatar contradiction clause 277]
283. r2_hidden(k4_tarski(X1,X2),X0) | ~r2_hidden(X2,sK1) | ~r2_hidden(X1,sK0) | ~v1_xboole_0(X0) <- (1) [superposition 216,58] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
286. ~r2_hidden(X2,sK1) | ~r2_hidden(X1,sK0) | ~v1_xboole_0(X0) <- (1) [subsumption resolution 283,62] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
287. 10 | 17 | 18 | ~1 [avatar split clause 286,83,265,262,191]
289. k1_xboole_0 = sK1 <- (18) [resolution 266,59] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
295. $false <- (~2, 18) [subsumption resolution 289,89] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
296. 2 | ~18 [avatar contradiction clause 295]
336. $false <- (8, 10) [unit resulting resolution 120,192] {a:4,w:0,sine:255,thAx:0,allAx:14,thDist:-14}
341. ~8 | ~10 [avatar contradiction clause 336]
414. $false <- (11, ~13) [unit resulting resolution 206,195,195,67] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:30,thDist:-30}
422. ~11 | 13 [avatar contradiction clause 414]
423. $false [avatar sat refutation 90,95,96,106,121,127,188,197,207,278,287,296,341,422] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:370,thDist:-370}
% SZS output end Proof for zfmisc_1__t90_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% Introduced skolems: 10
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 310
% Active clauses: 68
% Passive clauses: 148
% Final active clauses: 36
% Final passive clauses: 63
% Main loop iterations started: 79
% 
% >>> Simplifying Inferences
% Duplicate literals: 20
% Trivial inequalities: 4
% Fw subsumption resolutions: 8
% Fw demodulations: 6
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 1
% Forward subsumptions: 78
% Backward subsumptions: 19
% 
% >>> Generating Inferences
% Binary resolution: 77
% Unit resulting resolution: 6
% Forward superposition: 48
% Backward superposition: 60
% Self superposition: 3
% 
% >>> AVATAR
% Split clauses: 13
% Split components: 30
% Sat splitting refutations: 11
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 42
% SAT solver unit clauses: 5
% SAT solver binary clauses: 21
% 
% Memory used [KB]: 5117
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.005 s ( own 0.003 s ) 

