i: [1,221,0,2,1,1,0,0] t124_xreal_1
a: [3,356,0,2,1,221]
s: 356
i: [1,220,0,2,1,1,0,0] t124_xreal_1
a: [3,361,0,2,1,220]
s: 361
i: [1,223,0,3,1,1,0,0] t124_xreal_1
a: [3,346,0,3,1,223]
s: 346
i: [1,222,0,3,1,1,0,0] t124_xreal_1
a: [3,351,0,3,1,222]
s: 351
i: [1,224,0,5,1,1,0,0] t124_xreal_1
a: [3,341,0,5,1,224]
s: 341
i: [1,225,0,2,1,0,0,2] fc1_numbers
a: [3,366,0,2,1,225]
s: 366
i: [1,226,0,2,1,0,0,2] fc1_xboole_0
a: [3,371,0,2,1,226]
s: 371
i: [1,233,0,2,1,0,0,3] fc6_ordinal1
a: [3,401,0,2,1,233]
s: 401
i: [1,232,0,2,1,0,0,3] fc6_ordinal1
a: [3,406,0,2,1,232]
s: 406
i: [1,305,0,2,1,0,0,2] rc2_xboole_0
a: [3,418,0,2,1,305]
s: 418
i: [1,306,0,2,1,0,0,4] rc4_ordinal1
a: [3,423,0,2,1,306]
s: 423
i: [1,307,0,2,1,0,0,4] rc1_ordinal1
a: [3,428,0,2,1,307]
s: 428
i: [1,308,0,2,1,0,0,1] rc1_xreal_0
a: [3,433,0,2,1,308]
s: 433
i: [1,309,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,438,0,2,1,309]
s: 438
i: [1,310,0,2,1,0,0,2] rc1_xxreal_0
a: [3,443,0,2,1,310]
s: 443
i: [1,311,0,2,1,0,0,2] rc1_xboole_0
a: [3,448,0,2,1,311]
s: 448
i: [1,315,0,2,1,0,0,255] rc2_ordinal1
a: [3,453,0,2,1,315]
s: 453
i: [1,314,0,2,1,0,0,255] rc2_ordinal1
a: [3,458,0,2,1,314]
s: 458
i: [1,313,0,2,1,0,0,255] rc2_ordinal1
a: [3,463,0,2,1,313]
s: 463
i: [1,312,0,2,1,0,0,255] rc2_ordinal1
a: [3,468,0,2,1,312]
s: 468
i: [1,317,0,2,1,0,0,4] rc5_ordinal1
a: [3,473,0,2,1,317]
s: 473
i: [1,316,0,2,1,0,0,4] rc5_ordinal1
a: [3,478,0,2,1,316]
s: 478
i: [1,319,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,483,0,2,1,319]
s: 483
i: [1,318,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,488,0,2,1,318]
s: 488
i: [1,323,0,2,1,0,0,2] rc2_xreal_0
a: [3,493,0,2,1,323]
s: 493
i: [1,322,0,2,1,0,0,2] rc2_xreal_0
a: [3,498,0,2,1,322]
s: 498
i: [1,321,0,2,1,0,0,2] rc2_xreal_0
a: [3,503,0,2,1,321]
s: 503
i: [1,320,0,2,1,0,0,2] rc2_xreal_0
a: [3,508,0,2,1,320]
s: 508
i: [1,327,0,2,1,0,0,2] rc3_xreal_0
a: [3,513,0,2,1,327]
s: 513
i: [1,326,0,2,1,0,0,2] rc3_xreal_0
a: [3,518,0,2,1,326]
s: 518
i: [1,325,0,2,1,0,0,2] rc3_xreal_0
a: [3,523,0,2,1,325]
s: 523
i: [1,324,0,2,1,0,0,2] rc3_xreal_0
a: [3,528,0,2,1,324]
s: 528
i: [1,329,0,2,1,0,0,2] rc2_xxreal_0
a: [3,533,0,2,1,329]
s: 533
i: [1,328,0,2,1,0,0,2] rc2_xxreal_0
a: [3,538,0,2,1,328]
s: 538
i: [1,331,0,2,1,0,0,2] rc3_xxreal_0
a: [3,543,0,2,1,331]
s: 543
i: [1,330,0,2,1,0,0,2] rc3_xxreal_0
a: [3,548,0,2,1,330]
s: 548
i: [1,335,0,2,1,0,0,2] rc4_xreal_0
a: [3,553,0,2,1,335]
s: 553
i: [1,334,0,2,1,0,0,2] rc4_xreal_0
a: [3,558,0,2,1,334]
s: 558
i: [1,333,0,2,1,0,0,2] rc4_xreal_0
a: [3,563,0,2,1,333]
s: 563
i: [1,332,0,2,1,0,0,2] rc4_xreal_0
a: [3,568,0,2,1,332]
s: 568
i: [1,337,0,2,1,0,0,2] rc4_xxreal_0
a: [3,573,0,2,1,337]
s: 573
i: [1,336,0,2,1,0,0,2] rc4_xxreal_0
a: [3,578,0,2,1,336]
s: 578
i: [1,297,0,3,1,0,0,255] reflexivity_r1_tarski
s: 297
i: [1,227,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,376,0,3,1,227]
s: 376
r: 371
d: [2,580,0,2,1,2,23,371,376]
a: [3,584,0,2,1,580]
s: 584
i: [1,228,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,381,0,3,1,228]
s: 381
i: [1,229,0,3,1,0,0,2] t1_numerals
a: [3,386,0,3,1,229]
s: 386
d: [2,586,0,3,1,2,22,386,376]
a: [3,590,0,3,1,586]
s: 590
i: [1,235,0,4,2,0,0,4] cc6_ordinal1
s: 235
i: [1,236,0,4,2,0,0,4] cc2_xreal_0
s: 236
i: [1,237,0,4,2,0,0,4] cc2_xcmplx_0
s: 237
d: [2,592,1,2,1,1,40,235,423]
a: [3,597,1,2,1,592]
s: 597
i: [1,238,0,4,2,0,0,4] cc2_xxreal_0
s: 238
d: [2,593,1,2,1,1,40,235,473]
a: [3,602,1,2,1,593]
s: 602
i: [1,239,0,4,2,0,0,255] cc1_ordinal1
s: 239
i: [1,240,0,4,2,0,0,255] cc1_ordinal1
s: 240
d: [2,604,1,2,1,1,40,236,423]
a: [3,609,1,2,1,604]
s: 609
i: [1,242,0,4,2,0,0,2] cc3_xreal_0
s: 242
d: [2,690,1,2,1,1,40,242,361]
a: [3,700,1,2,1,690]
s: 700
i: [1,243,0,4,2,0,0,2] cc4_xreal_0
s: 243
i: [1,256,0,4,2,0,0,255] cc9_ordinal1
s: 256
d: [2,691,1,2,1,1,40,242,356]
a: [3,705,1,2,1,691]
s: 705
i: [1,257,0,4,2,0,0,4] cc7_ordinal1
s: 257
d: [2,713,1,2,1,1,40,243,361]
a: [3,723,1,2,1,713]
s: 723
i: [1,258,0,4,2,0,0,4] cc3_ordinal1
s: 258
i: [1,296,0,4,1,0,0,3] existence_m1_subset_1
s: 296
d: [2,714,1,2,1,1,40,243,356]
a: [3,728,1,2,1,714]
s: 728
i: [1,230,0,4,1,0,0,2] dt_k5_numbers
a: [3,391,0,4,1,230]
s: 391
d: [2,605,1,2,1,1,40,236,473]
a: [3,614,1,2,1,605]
s: 614
i: [1,231,0,4,1,0,0,1] dt_k6_numbers
a: [3,396,0,4,1,231]
s: 396
d: [2,808,0,4,1,2,22,391,381]
a: [3,812,0,4,1,808]
s: 812
d: [2,616,1,2,1,1,40,237,423]
a: [3,621,1,2,1,616]
s: 621
d: [2,818,0,4,1,2,22,396,381]
a: [3,822,0,4,1,818]
s: 822
d: [2,617,1,2,1,1,40,237,473]
a: [3,626,1,2,1,617]
s: 626
i: [1,244,0,5,2,0,0,1] fc4_xreal_0
s: 244
i: [1,245,0,5,2,0,0,1] fc4_xreal_0
s: 245
d: [2,628,1,2,1,1,40,238,423]
a: [3,633,1,2,1,628]
s: 633
i: [1,254,0,5,2,0,0,1] dt_k5_xcmplx_0
s: 254
d: [2,629,1,2,1,1,40,238,473]
a: [3,638,1,2,1,629]
s: 638
i: [1,259,0,5,2,0,0,2] t6_boole
s: 259
i: [1,260,0,5,2,0,0,3] cc8_ordinal1
s: 260
d: [2,640,1,2,1,1,40,239,401]
a: [3,648,1,2,1,640]
s: 648
i: [1,261,0,5,2,0,0,2] cc1_xreal_0
s: 261
d: [2,642,1,2,1,1,40,239,428]
a: [3,658,1,2,1,642]
s: 658
i: [1,262,0,5,2,0,0,2] cc1_xcmplx_0
s: 262
i: [1,293,0,5,2,0,0,255] t7_boole
s: 293
d: [2,665,1,2,1,1,40,240,401]
a: [3,673,1,2,1,665]
s: 673
i: [1,303,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,413,0,5,2,303]
s: 413
d: [2,667,1,2,1,1,40,240,428]
a: [3,683,1,2,1,667]
s: 683
d: [2,826,0,5,2,1,22,259,376]
s: 826
i: [1,273,0,6,3,0,0,255] cc2_ordinal1
s: 273
d: [2,851,1,5,2,2,44,351,826]
s: 851
i: [1,280,0,6,3,0,0,2] cc5_xxreal_0
s: 280
d: [2,856,2,4,2,3,40,851,413]
d: [2,857,2,2,1,4,21,856,723]
a: [3,861,2,2,1,857]
s: 861
i: [1,282,0,6,3,0,0,2] cc5_xxreal_0
s: 282
i: [1,283,0,6,3,0,0,2] cc3_xxreal_0
s: 283
d: [2,852,1,5,2,2,44,346,826]
s: 852
i: [1,285,0,6,3,0,0,2] cc3_xxreal_0
s: 285
d: [2,905,2,4,2,3,40,852,413]
d: [2,906,2,2,1,4,21,905,728]
a: [3,910,2,2,1,906]
s: 910
i: [1,287,0,6,3,0,0,2] cc7_xxreal_0
s: 287
i: [1,288,0,6,3,0,0,2] cc7_xxreal_0
s: 288
d: [2,693,1,2,1,1,40,242,433]
a: [3,711,1,2,1,693]
s: 711
i: [1,289,0,6,2,0,0,255] t1_subset
s: 289
d: [2,716,1,2,1,1,40,243,433]
a: [3,734,1,2,1,716]
s: 734
i: [1,298,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 298
i: [1,241,0,7,3,0,0,4] cc5_ordinal1
s: 241
d: [2,736,1,2,1,1,40,256,584]
a: [3,743,1,2,1,736]
s: 743
i: [1,255,0,7,2,0,0,2] involutiveness_k5_xcmplx_0
s: 255
d: [2,737,1,2,1,1,40,256,448]
a: [3,748,1,2,1,737]
s: 748
i: [1,263,0,7,3,0,0,1] fc28_xreal_0
s: 263
i: [1,264,0,7,3,0,0,1] fc28_xreal_0
s: 264
d: [2,738,1,2,1,1,40,256,568]
a: [3,753,1,2,1,738]
s: 753
i: [1,265,0,7,3,0,0,1] fc30_xreal_0
s: 265
d: [2,739,1,2,1,1,40,256,578]
a: [3,758,1,2,1,739]
s: 758
i: [1,266,0,7,3,0,0,1] fc30_xreal_0
s: 266
i: [1,267,0,7,3,0,0,1] fc7_xcmplx_0
s: 267
d: [2,760,1,2,1,1,40,257,584]
a: [3,767,1,2,1,760]
s: 767
i: [1,268,0,7,3,0,0,1] fc7_xcmplx_0
s: 268
d: [2,761,1,2,1,1,40,257,448]
a: [3,772,1,2,1,761]
s: 772
i: [1,274,0,7,3,0,0,1] fc27_xreal_0
s: 274
i: [1,275,0,7,3,0,0,1] fc27_xreal_0
s: 275
d: [2,762,1,2,1,1,40,257,568]
a: [3,777,1,2,1,762]
s: 777
i: [1,276,0,7,3,0,0,1] fc29_xreal_0
s: 276
d: [2,763,1,2,1,1,40,257,578]
a: [3,782,1,2,1,763]
s: 782
i: [1,277,0,7,3,0,0,1] fc29_xreal_0
s: 277
i: [1,291,0,7,2,0,0,255] t3_subset
s: 291
d: [2,784,1,2,1,1,40,258,584]
a: [3,791,1,2,1,784]
s: 791
i: [1,292,0,7,3,0,0,2] t8_boole
s: 292
d: [2,785,1,2,1,1,40,258,448]
a: [3,796,1,2,1,785]
s: 796
i: [1,234,0,8,3,0,0,255] cc10_ordinal1
s: 234
i: [1,270,0,8,4,0,0,2] cc4_xxreal_0
s: 270
r: 413
r: 910
a: [3,909,2,2,1,906]
s: 909
r: 909
a: [3,410,0,2,1,303]
s: 410
r: 280
r: 282
r: 283
r: 270
d: [2,1147,1,0,0,2,40,410,573]
e: 1147
d: [2,1146,1,0,0,2,40,410,558]
e: 1146
d: [2,1145,1,0,0,2,40,410,548]
e: 1145
d: [2,1144,1,0,0,2,40,410,538]
e: 1144
d: [2,1143,1,0,0,2,40,410,523]
e: 1143
d: [2,1142,1,0,0,2,40,410,503]
e: 1142
d: [2,1141,2,0,0,2,40,410,638]
e: 1141
d: [2,1140,1,0,0,2,40,410,443]
e: 1140
d: [2,1139,2,0,0,2,40,410,734]
e: 1139
d: [2,1138,2,0,0,2,40,410,633]
e: 1138
d: [2,1137,2,0,0,2,40,410,728]
e: 1137
d: [2,1136,2,0,0,2,40,410,723]
e: 1136
d: [2,1133,2,0,0,2,52,723,410]
e: 1133
d: [2,1134,1,0,0,2,52,361,410,243]
e: 1134
d: [2,1135,2,0,0,2,52,767,410,238]
e: 1135
r: 410
s: 910
s: 413
s: 280
d: [2,1033,2,6,3,1,40,270,723]
d: [2,1045,2,4,2,2,21,1033,861]
a: [3,1049,2,2,1,1045]
s: 1049
s: 282
d: [2,1034,2,6,3,1,40,270,728]
d: [2,1055,2,4,2,2,21,1034,910]
a: [3,1059,2,2,1,1055]
s: 1059
s: 283
s: 270
d: [2,786,1,2,1,1,40,258,568]
a: [3,801,1,2,1,786]
s: 801
i: [1,272,0,8,4,0,0,2] cc6_xxreal_0
s: 272
d: [2,787,1,2,1,1,40,258,578]
a: [3,806,1,2,1,787]
s: 806
i: [1,278,0,8,4,0,0,2] cc8_xxreal_0
s: 278
i: [1,290,0,8,3,0,0,255] t2_subset
s: 290
d: [2,863,1,4,2,1,40,280,518]
d: [2,865,1,2,1,2,21,863,523]
a: [3,869,1,2,1,865]
s: 869
i: [1,295,0,9,3,0,0,255] t5_subset
s: 295
d: [2,864,1,4,2,1,40,280,543]
d: [2,871,1,2,1,2,21,864,548]
a: [3,875,1,2,1,871]
s: 875
i: [1,294,0,10,3,0,0,255] t4_subset
s: 294
i: [1,304,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 304
d: [2,1254,1,7,3,1,40,304,346]
d: [2,1260,1,5,2,2,21,1254,728]
a: [3,1264,1,3,1,1260]
s: 1264
i: [1,248,0,11,5,0,0,1] t2_real
s: 248
a: [3,1062,2,2,1,1055]
s: 1062
i: [1,249,0,11,5,0,0,1] t1_real
s: 249
i: [1,250,0,11,5,0,0,1] t7_real
s: 250
a: [3,1053,2,2,1,1045]
s: 1053
d: [2,1340,3,0,0,3,52,723,1049,1053,282]
e: 1340
d: [2,1341,3,4,2,1,40,1053,282]
d: [2,1344,3,2,1,2,21,1341,1049]
d: [2,1345,3,0,0,3,21,1344,723]
e: 1345
r: 1053
i: [1,251,0,11,5,0,0,1] t8_real
s: 251
d: [2,1323,1,8,4,1,40,250,351]
d: [2,1329,1,6,3,2,21,1323,361]
a: [3,1333,1,2,1,1329]
s: 1333
r: 1049
r: 1333
i: [1,252,0,11,5,0,0,1] t4_real
s: 252
i: [1,253,0,11,5,0,0,1] t3_real
s: 253
a: [3,1332,1,2,1,1329]
s: 1332
i: [1,246,0,13,6,0,0,1] t6_real
s: 246
a: [3,1048,2,2,1,1045]
s: 1048
i: [1,247,0,13,6,0,0,1] t5_real
s: 247
i: [1,302,0,14,4,0,0,2] existence_m2_subset_1
s: 302
s: 1053
i: [1,299,0,15,5,0,0,2] dt_m2_subset_1
s: 299
d: [2,1256,1,7,3,1,40,304,351]
d: [2,1267,1,5,2,2,21,1256,723]
a: [3,1271,1,3,1,1267]
s: 1271
i: [1,300,0,15,5,0,0,2] redefinition_m2_subset_1
s: 300
i: [1,301,0,15,5,0,0,2] redefinition_m2_subset_1
s: 301
d: [2,1275,1,12,4,1,40,248,341]
a: [3,1283,1,3,1,1275]
s: 1283
d: [2,1412,2,0,0,2,52,356,1283,245]
e: 1412
d: [2,1413,2,2,1,1,40,1283,245]
d: [2,1415,2,0,0,2,21,1413,356]
e: 1415
r: 1283
a: [3,1282,1,3,1,1275]
s: 1282
a: [3,1287,1,3,1,1275]
s: 1287
d: [2,1429,2,0,0,2,52,361,1287,245]
e: 1429
d: [2,1430,2,2,1,1,40,1287,245]
d: [2,1432,2,0,0,2,21,1430,361]
e: 1432
r: 1287
a: [3,1286,1,3,1,1275]
s: 1286
a: [3,1291,1,3,1,1275]
s: 1291
d: [2,1446,3,0,0,3,52,1062,356,1291,266]
e: 1446
d: [2,1447,2,4,2,1,40,1291,266]
d: [2,1451,2,2,1,2,21,1447,356]
d: [2,1452,2,0,0,3,21,1451,1062]
e: 1452
r: 1062
a: [3,1063,2,2,1,1055]
s: 1063
d: [2,1457,3,0,0,3,52,728,1059,1063,282]
e: 1457
d: [2,1458,3,4,2,1,40,1063,282]
d: [2,1461,3,2,1,2,21,1458,1059]
d: [2,1462,3,0,0,3,21,1461,728]
e: 1462
r: 1059
r: 1063
r: 1291
a: [3,1290,1,3,1,1275]
s: 1290
s: 1059
a: [3,1295,1,3,1,1275]
s: 1295
d: [2,1482,3,0,0,3,52,1053,361,1295,277]
e: 1482
d: [2,1483,2,4,2,1,40,1295,277]
d: [2,1485,2,2,1,2,21,1483,361]
d: [2,1486,2,0,0,3,21,1485,1053]
e: 1486
r: 1059
r: 1290
r: 1295
s: 1291
a: [3,1058,2,2,1,1055]
s: 1058
a: [3,1294,1,3,1,1275]
s: 1294
s: 1063
d: [2,1305,1,12,4,1,40,249,341]
a: [3,1312,1,3,1,1305]
s: 1312
a: [3,1317,1,3,1,1305]
s: 1317
d: [2,1278,2,8,4,1,40,248,1264]
d: [2,1299,2,6,3,2,21,1278,356]
a: [3,1303,2,2,1,1299]
s: 1303
d: [2,827,1,2,1,1,40,260,590]
d: [2,925,2,2,1,1,40,767,238]
d: [2,927,2,2,1,1,40,767,236]
d: [2,1347,2,4,2,1,40,1333,283]
d: [2,1349,2,2,1,2,21,1347,584]
d: [2,1373,1,15,5,1,40,247,341]
d: [2,1434,2,3,1,1,40,1286,243]
d: [2,1464,3,2,1,2,21,1458,728]
d: [2,1321,1,8,4,1,40,250,346]
d: [2,1474,1,6,3,2,21,1321,356]
d: [2,1488,2,4,2,1,40,1291,266]
d: [2,1491,2,2,1,2,21,1488,356]
d: [2,1496,2,6,2,1,40,1294,282]
d: [2,1497,2,6,2,1,40,1294,280]
d: [2,1505,3,4,2,1,40,1303,280]
d: [2,1507,3,2,1,2,21,1505,584]
f: 1507,1497,1496,1491,1486,1474,1464,1434,1432,1415,1373,1349,1329,927,925,827,714,586,580,229,227,226,220,221,222,223,224
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t124_xreal_1
% SZS output start Proof for xreal_1__t124_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k5_xcmplx_0(X1),k5_xcmplx_0(X0)) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(k6_numbers,X0)))) [input t124_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k5_xcmplx_0(X1),k5_xcmplx_0(X0)) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(k6_numbers,X0)))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
14. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
19. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
40. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
43. ! [X0] : ((v1_xreal_0(X0) & v3_xxreal_0(X0)) => (v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0)))) [input fc29_xreal_0]
44. ! [X0] : ((v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => (~v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0)))) [input fc30_xreal_0]
45. ! [X0] : (v1_xreal_0(X0) => (v1_xreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0)))) [input fc4_xreal_0]
67. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
71. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
80. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v2_xxreal_0(X1) | v3_xxreal_0(X0) | v1_xboole_0(X1))))) [input t5_real]
85. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v3_xxreal_0(X1) & ~v2_xxreal_0(X0) & ~r1_xxreal_0(X0,X1)))) [input t7_real]
99. ? [X0] : (? [X1] : ((r1_xxreal_0(k5_xcmplx_0(X1),k5_xcmplx_0(X0)) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(k6_numbers,X0)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
100. ? [X0] : (? [X1] : (r1_xxreal_0(k5_xcmplx_0(X1),k5_xcmplx_0(X0)) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 99]
103. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
105. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
109. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
110. ! [X0] : ((v1_xreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | ~v1_xreal_0(X0)) [ennf transformation 45]
113. ! [X0] : (! [X1] : (((v2_xxreal_0(X1) | v3_xxreal_0(X0) | v1_xboole_0(X1)) | ~r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 80]
114. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | v3_xxreal_0(X0) | v1_xboole_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 113]
119. ! [X0] : (! [X1] : ((v3_xxreal_0(X1) | v2_xxreal_0(X0) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 85]
120. ! [X0] : (! [X1] : (v3_xxreal_0(X1) | v2_xxreal_0(X0) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 119]
133. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
138. ! [X0] : ((~v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | (~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 44]
139. ! [X0] : ((~v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 138]
150. ! [X0] : ((v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | (~v1_xreal_0(X0) | ~v3_xxreal_0(X0))) [ennf transformation 43]
151. ! [X0] : ((v3_xxreal_0(k5_xcmplx_0(X0)) & v1_xcmplx_0(k5_xcmplx_0(X0))) | ~v1_xreal_0(X0) | ~v3_xxreal_0(X0)) [flattening 150]
154. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 19]
155. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 154]
156. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 14]
157. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 156]
180. ? [X0] : (? [X1] : (r1_xxreal_0(k5_xcmplx_0(X1),k5_xcmplx_0(X0)) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (r1_xxreal_0(k5_xcmplx_0(X1),k5_xcmplx_0(sK0)) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
181. ? [X1] : (r1_xxreal_0(k5_xcmplx_0(X1),k5_xcmplx_0(sK0)) & ~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) => (r1_xxreal_0(k5_xcmplx_0(sK1),k5_xcmplx_0(sK0)) & ~r1_xxreal_0(sK1,k6_numbers) & ~r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) [choice axiom]
182. (r1_xxreal_0(k5_xcmplx_0(sK1),k5_xcmplx_0(sK0)) & ~r1_xxreal_0(sK1,k6_numbers) & ~r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 100,181,180]
220. v1_xreal_0(sK0) [cnf transformation 182] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
221. v1_xreal_0(sK1) [cnf transformation 182] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
222. ~r1_xxreal_0(k6_numbers,sK0) [cnf transformation 182] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
223. ~r1_xxreal_0(sK1,k6_numbers) [cnf transformation 182] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
224. r1_xxreal_0(k5_xcmplx_0(sK1),k5_xcmplx_0(sK0)) [cnf transformation 182] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
226. v1_xboole_0(k1_xboole_0) [cnf transformation 40] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
227. k6_numbers = k1_xboole_0 [cnf transformation 67] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
229. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 71] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
236. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 103] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
238. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 105] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
243. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 109] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
245. v1_xreal_0(k5_xcmplx_0(X0)) | ~v1_xreal_0(X0) [cnf transformation 110] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
247. ~r1_xxreal_0(X0,X1) | v3_xxreal_0(X0) | v1_xboole_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 114] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
250. r1_xxreal_0(X0,X1) | v2_xxreal_0(X0) | v3_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 120] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
260. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 133] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
266. ~v3_xxreal_0(k5_xcmplx_0(X0)) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 139] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
277. v3_xxreal_0(k5_xcmplx_0(X0)) | ~v1_xreal_0(X0) | ~v3_xxreal_0(X0) [cnf transformation 151] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
280. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 155] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
282. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 155] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
283. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 157] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
339. 1 <=> r1_xxreal_0(k5_xcmplx_0(sK1),k5_xcmplx_0(sK0)) [avatar definition]
341. r1_xxreal_0(k5_xcmplx_0(sK1),k5_xcmplx_0(sK0)) <- (1) [avatar component clause 339] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
342. 1 [avatar split clause 224,339]
344. 2 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
346. ~r1_xxreal_0(sK1,k6_numbers) <- (~2) [avatar component clause 344] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
347. ~2 [avatar split clause 223,344]
349. 3 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
351. ~r1_xxreal_0(k6_numbers,sK0) <- (~3) [avatar component clause 349] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
352. ~3 [avatar split clause 222,349]
354. 4 <=> v1_xreal_0(sK1) [avatar definition]
356. v1_xreal_0(sK1) <- (4) [avatar component clause 354] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
357. 4 [avatar split clause 221,354]
359. 5 <=> v1_xreal_0(sK0) [avatar definition]
361. v1_xreal_0(sK0) <- (5) [avatar component clause 359] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
362. 5 [avatar split clause 220,359]
369. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
371. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 369] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
372. 7 [avatar split clause 226,369]
374. 8 <=> k6_numbers = k1_xboole_0 [avatar definition]
376. k6_numbers = k1_xboole_0 <- (8) [avatar component clause 374] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
377. 8 [avatar split clause 227,374]
384. 10 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
386. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (10) [avatar component clause 384] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
387. 10 [avatar split clause 229,384]
580. v1_xboole_0(k6_numbers) <- (7, 8) [backward demodulation 371,376] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
582. 50 <=> v1_xboole_0(k6_numbers) [avatar definition]
584. v1_xboole_0(k6_numbers) <- (50) [avatar component clause 582] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
585. 50 | ~7 | ~8 [avatar split clause 580,374,369,582]
586. m1_subset_1(k6_numbers,k4_ordinal1) <- (8, 10) [forward demodulation 386,376] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
588. 51 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
590. m1_subset_1(k6_numbers,k4_ordinal1) <- (51) [avatar component clause 588] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
591. 51 | ~8 | ~10 [avatar split clause 586,384,374,588]
714. v1_xxreal_0(sK1) <- (4) [resolution 243,356] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
726. 72 <=> v1_xxreal_0(sK1) [avatar definition]
728. v1_xxreal_0(sK1) <- (72) [avatar component clause 726] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
729. 72 | ~4 [avatar split clause 714,354,726]
765. 78 <=> v7_ordinal1(k6_numbers) [avatar definition]
767. v7_ordinal1(k6_numbers) <- (78) [avatar component clause 765] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
827. v7_ordinal1(k6_numbers) <- (51) [resolution 260,590] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
829. 78 | ~51 [avatar split clause 827,588,765]
925. v1_xxreal_0(k6_numbers) <- (78) [resolution 767,238] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
927. v1_xreal_0(k6_numbers) <- (78) [resolution 767,236] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
931. 99 <=> v1_xxreal_0(k6_numbers) [avatar definition]
934. 99 | ~78 [avatar split clause 925,765,931]
941. 101 <=> v1_xreal_0(k6_numbers) [avatar definition]
944. 101 | ~78 [avatar split clause 927,765,941]
1051. 114 <=> v3_xxreal_0(sK0) [avatar definition]
1053. v3_xxreal_0(sK0) <- (114) [avatar component clause 1051] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1057. 115 <=> v2_xxreal_0(sK1) [avatar definition]
1061. 116 <=> v3_xxreal_0(sK1) [avatar definition]
1063. v3_xxreal_0(sK1) <- (116) [avatar component clause 1061] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1281. 140 <=> v1_xreal_0(k5_xcmplx_0(sK1)) [avatar definition]
1283. ~v1_xreal_0(k5_xcmplx_0(sK1)) <- (~140) [avatar component clause 1281] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1285. 141 <=> v1_xreal_0(k5_xcmplx_0(sK0)) [avatar definition]
1286. v1_xreal_0(k5_xcmplx_0(sK0)) <- (141) [avatar component clause 1285] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1287. ~v1_xreal_0(k5_xcmplx_0(sK0)) <- (~141) [avatar component clause 1285] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1289. 142 <=> v3_xxreal_0(k5_xcmplx_0(sK1)) [avatar definition]
1291. v3_xxreal_0(k5_xcmplx_0(sK1)) <- (142) [avatar component clause 1289] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1293. 143 <=> v3_xxreal_0(k5_xcmplx_0(sK0)) [avatar definition]
1294. v3_xxreal_0(k5_xcmplx_0(sK0)) <- (143) [avatar component clause 1293] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1295. ~v3_xxreal_0(k5_xcmplx_0(sK0)) <- (~143) [avatar component clause 1293] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1301. 144 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1303. v3_xxreal_0(k6_numbers) <- (144) [avatar component clause 1301] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1311. 145 <=> v2_xxreal_0(k5_xcmplx_0(sK0)) [avatar definition]
1321. v2_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (~2) [resolution 250,346] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1323. v2_xxreal_0(k6_numbers) | v3_xxreal_0(sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (~3) [resolution 250,351] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1329. v2_xxreal_0(k6_numbers) | v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (~3, 5) [subsumption resolution 1323,361] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1331. 147 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1333. v2_xxreal_0(k6_numbers) <- (147) [avatar component clause 1331] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1334. ~101 | 114 | 147 | 3 | ~5 [avatar split clause 1329,359,349,1331,1051,941]
1347. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (147) [resolution 1333,283] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1349. ~v1_xxreal_0(k6_numbers) <- (50, 147) [subsumption resolution 1347,584] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1350. ~99 | ~50 | ~147 [avatar split clause 1349,1331,582,931]
1373. v3_xxreal_0(k5_xcmplx_0(sK1)) | v1_xboole_0(k5_xcmplx_0(sK0)) | v2_xxreal_0(k5_xcmplx_0(sK0)) | ~v1_xreal_0(k5_xcmplx_0(sK0)) | ~v1_xreal_0(k5_xcmplx_0(sK1)) <- (1) [resolution 247,341] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1381. 152 <=> v1_xboole_0(k5_xcmplx_0(sK0)) [avatar definition]
1384. ~140 | ~141 | 145 | 152 | 142 | ~1 [avatar split clause 1373,339,1289,1381,1311,1285,1281]
1413. ~v1_xreal_0(sK1) <- (~140) [resolution 1283,245] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1415. $false <- (4, ~140) [subsumption resolution 1413,356] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1416. ~4 | 140 [avatar contradiction clause 1415]
1430. ~v1_xreal_0(sK0) <- (~141) [resolution 1287,245] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1432. $false <- (5, ~141) [subsumption resolution 1430,361] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1433. ~5 | 141 [avatar contradiction clause 1432]
1434. v1_xxreal_0(k5_xcmplx_0(sK0)) <- (141) [resolution 1286,243] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1437. 156 <=> v1_xxreal_0(k5_xcmplx_0(sK0)) [avatar definition]
1440. 156 | ~141 [avatar split clause 1434,1285,1437]
1458. ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (116) [resolution 1063,282] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1464. ~v2_xxreal_0(sK1) <- (72, 116) [subsumption resolution 1458,728] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1465. ~115 | ~72 | ~116 [avatar split clause 1464,1061,726,1057]
1474. v2_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (~2, 4) [subsumption resolution 1321,356] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1475. ~101 | 144 | 115 | 2 | ~4 [avatar split clause 1474,354,344,1057,1301,941]
1483. ~v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) <- (~143) [resolution 1295,277] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1485. ~v3_xxreal_0(sK0) <- (5, ~143) [subsumption resolution 1483,361] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1486. $false <- (5, 114, ~143) [subsumption resolution 1485,1053] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1487. ~5 | ~114 | 143 [avatar contradiction clause 1486]
1488. ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) <- (142) [resolution 1291,266] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1491. v3_xxreal_0(sK1) <- (4, 142) [subsumption resolution 1488,356] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1492. 116 | ~4 | ~142 [avatar split clause 1491,1289,354,1061]
1496. ~v2_xxreal_0(k5_xcmplx_0(sK0)) | ~v1_xxreal_0(k5_xcmplx_0(sK0)) <- (143) [resolution 1294,282] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1497. ~v1_xboole_0(k5_xcmplx_0(sK0)) | ~v1_xxreal_0(k5_xcmplx_0(sK0)) <- (143) [resolution 1294,280] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1498. ~156 | ~145 | ~143 [avatar split clause 1496,1293,1311,1437]
1499. ~156 | ~152 | ~143 [avatar split clause 1497,1293,1381,1437]
1505. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (144) [resolution 1303,280] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1507. ~v1_xxreal_0(k6_numbers) <- (50, 144) [subsumption resolution 1505,584] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1508. ~99 | ~50 | ~144 [avatar split clause 1507,1301,582,931]
1509. $false [avatar sat refutation 342,347,352,357,362,372,377,387,585,591,729,829,934,944,1334,1350,1384,1416,1433,1440,1465,1475,1487,1492,1498,1499,1508] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1141,thDist:-1141}
% SZS output end Proof for xreal_1__t124_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 86
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 118
% Generated clauses: 636
% Active clauses: 170
% Passive clauses: 274
% Final active clauses: 149
% Final passive clauses: 84
% Main loop iterations started: 187
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 58
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 81
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 231
% Unit resulting resolution: 9
% Factoring: 1
% Forward superposition: 13
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 36
% Split components: 93
% Sat splitting refutations: 27
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 208
% SAT solver unit clauses: 50
% SAT solver binary clauses: 88
% 
% Memory used [KB]: 5756
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% congruence closure: 0.003 s
% other: 0.008 s ( own 0.003 s ) 

