i: [1,39,0,2,1,1,0,0] t175_relat_1
a: [3,74,0,2,1,39]
s: 74
i: [1,38,0,2,1,1,0,0] t175_relat_1
a: [3,79,0,2,1,38]
s: 79
i: [1,41,0,3,1,1,0,0] t175_relat_1
a: [3,64,0,3,1,41]
s: 64
i: [1,40,0,3,1,1,0,0] t175_relat_1
a: [3,69,0,3,1,40]
s: 69
i: [1,42,0,5,1,1,0,0] t175_relat_1
a: [3,59,0,5,1,42]
s: 59
i: [1,48,0,3,1,0,0,1] reflexivity_r1_tarski
s: 48
i: [1,47,0,4,1,0,0,255] existence_m1_subset_1
s: 47
i: [1,49,0,6,2,0,0,1] dt_k5_relat_1
s: 49
i: [1,45,0,7,2,0,0,255] t3_subset
s: 45
i: [1,46,0,7,2,0,0,255] t3_subset
s: 46
i: [1,43,0,8,3,0,0,255] cc2_relat_1
s: 43
d: [2,81,1,5,1,0,40,45,47]
s: 81
i: [1,53,0,9,3,0,0,1] fc27_relat_1
s: 53
d: [2,83,1,6,2,0,40,43,47]
s: 83
i: [1,51,0,10,3,0,0,255] rd8_relat_1
s: 51
i: [1,54,0,10,3,0,0,1] fc27_relat_1
s: 54
d: [2,85,1,8,3,1,44,59,51]
d: [2,88,1,5,2,2,21,85,64]
d: [2,89,1,3,1,3,21,88,79]
a: [3,93,1,3,1,89]
s: 93
i: [1,55,0,10,3,0,0,1] fc27_relat_1
s: 55
d: [2,84,1,7,3,0,40,43,46]
s: 84
i: [1,50,0,11,2,0,0,255] rd5_relat_1
s: 50
i: [1,52,0,12,4,0,0,255] cc5_relat_1
s: 52
d: [2,109,1,10,3,0,40,52,47]
s: 109
i: [1,44,0,14,4,0,0,1] t76_relat_1
s: 44
d: [2,110,1,11,4,0,40,52,46]
s: 110
d: [2,111,0,12,3,0,21,44,84]
s: 111
d: [2,105,1,16,4,0,44,54,50]
d: [2,108,1,12,3,0,21,105,49]
s: 108
d: [2,112,2,8,3,1,40,110,64]
d: [2,115,2,6,2,2,21,112,79]
s: 115
d: [2,120,1,15,5,0,44,111,51]
d: [2,122,1,13,4,0,18,120]
s: 122
d: [2,117,2,13,4,0,40,111,84]
d: [2,124,2,9,3,0,21,117,49]
s: 124
d: [2,118,1,15,5,0,44,111,51]
d: [2,125,1,13,4,0,21,118,84]
s: 125
d: [2,142,2,0,0,4,52,79,69,64,59,125]
e: 142
d: [2,143,2,8,3,1,40,125,59]
d: [2,150,2,5,2,2,21,143,64]
d: [2,151,2,3,1,3,21,150,79]
d: [2,152,2,0,0,4,21,151,69]
e: 152
f: 152,38,40,41,42
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t175_relat_1
% SZS output start Proof for relat_1__t175_relat_1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : ((v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r1_tarski(X2,X1) => r1_tarski(X2,k5_relat_1(X1,X0))))) [input t175_relat_1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : ((v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r1_tarski(X2,X1) => r1_tarski(X2,k5_relat_1(X1,X0))))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
11. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => k5_relat_1(X1,X0) = X1) [input rd8_relat_1]
13. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
14. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0))))) [input t76_relat_1]
15. ~! [X0,X1] : (v1_relat_1(X1) => ! [X2] : ((v4_relat_1(X2,X0) & v1_relat_1(X2)) => (r1_tarski(X2,X1) => r1_tarski(X2,k5_relat_1(X1,X0))))) [flattening 2]
16. ! [X0,X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (r1_tarski(X1,X2) => r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0))))) [flattening 14]
17. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 13]
19. ? [X0,X1] : (? [X2] : ((~r1_tarski(X2,k5_relat_1(X1,X0)) & r1_tarski(X2,X1)) & (v4_relat_1(X2,X0) & v1_relat_1(X2))) & v1_relat_1(X1)) [ennf transformation 15]
20. ? [X0,X1] : (? [X2] : (~r1_tarski(X2,k5_relat_1(X1,X0)) & r1_tarski(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_relat_1(X1)) [flattening 19]
21. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
22. ! [X0,X1] : (! [X2] : ((r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [ennf transformation 16]
23. ! [X0,X1] : (! [X2] : (r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) [flattening 22]
26. ! [X0,X1] : (k5_relat_1(X1,X0) = X1 | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 11]
27. ! [X0,X1] : (k5_relat_1(X1,X0) = X1 | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 26]
32. ? [X0,X1] : (? [X2] : (~r1_tarski(X2,k5_relat_1(X1,X0)) & r1_tarski(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_relat_1(X1)) => (? [X2] : (~r1_tarski(X2,k5_relat_1(sK1,sK0)) & r1_tarski(X2,sK1) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_relat_1(sK1)) [choice axiom]
33. ? [X2] : (~r1_tarski(X2,k5_relat_1(sK1,sK0)) & r1_tarski(X2,sK1) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (~r1_tarski(sK2,k5_relat_1(sK1,sK0)) & r1_tarski(sK2,sK1) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
34. (~r1_tarski(sK2,k5_relat_1(sK1,sK0)) & r1_tarski(sK2,sK1) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_relat_1(sK1) [skolemisation 20,33,32]
35. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 17]
38. v1_relat_1(sK1) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. v4_relat_1(sK2,sK0) [cnf transformation 34] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. r1_tarski(sK2,sK1) [cnf transformation 34] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. ~r1_tarski(sK2,k5_relat_1(sK1,sK0)) [cnf transformation 34] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
43. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 21] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) [cnf transformation 23] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:1,thDist:-1}
46. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 35] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
51. k5_relat_1(X1,X0) = X1 | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 27] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
57. 1 <=> r1_tarski(sK2,k5_relat_1(sK1,sK0)) [avatar definition]
59. ~r1_tarski(sK2,k5_relat_1(sK1,sK0)) <- (~1) [avatar component clause 57] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
60. ~1 [avatar split clause 42,57]
62. 2 <=> r1_tarski(sK2,sK1) [avatar definition]
64. r1_tarski(sK2,sK1) <- (2) [avatar component clause 62] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
65. 2 [avatar split clause 41,62]
67. 3 <=> v4_relat_1(sK2,sK0) [avatar definition]
69. v4_relat_1(sK2,sK0) <- (3) [avatar component clause 67] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
70. 3 [avatar split clause 40,67]
77. 5 <=> v1_relat_1(sK1) [avatar definition]
79. v1_relat_1(sK1) <- (5) [avatar component clause 77] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
80. 5 [avatar split clause 38,77]
84. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 43,46] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
111. r1_tarski(k5_relat_1(X1,X0),k5_relat_1(X2,X0)) | ~r1_tarski(X1,X2) | ~v1_relat_1(X2) [subsumption resolution 44,84] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
118. r1_tarski(X0,k5_relat_1(X2,X1)) | ~r1_tarski(X0,X2) | ~v1_relat_1(X2) | ~v4_relat_1(X0,X1) | ~v1_relat_1(X0) [superposition 111,51] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
125. r1_tarski(X0,k5_relat_1(X2,X1)) | ~r1_tarski(X0,X2) | ~v1_relat_1(X2) | ~v4_relat_1(X0,X1) [subsumption resolution 118,84] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
143. ~r1_tarski(sK2,sK1) | ~v1_relat_1(sK1) | ~v4_relat_1(sK2,sK0) <- (~1) [resolution 125,59] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
150. ~v1_relat_1(sK1) | ~v4_relat_1(sK2,sK0) <- (~1, 2) [subsumption resolution 143,64] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
151. ~v4_relat_1(sK2,sK0) <- (~1, 2, 5) [subsumption resolution 150,79] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
152. $false <- (~1, 2, 3, 5) [subsumption resolution 151,69] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
153. 1 | ~2 | ~3 | ~5 [avatar contradiction clause 152]
155. $false [avatar sat refutation 60,65,70,80,153] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:47,thDist:-47}
% SZS output end Proof for relat_1__t175_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 18
% Generated clauses: 91
% Active clauses: 28
% Passive clauses: 35
% Final active clauses: 28
% Final passive clauses: 5
% Main loop iterations started: 30
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 21
% Unit resulting resolution: 1
% Forward superposition: 15
% Backward superposition: 5
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

