i: [1,233,0,2,1,1,0,0] t34_xreal_1
a: [3,365,0,2,1,233]
s: 365
i: [1,232,0,2,1,1,0,0] t34_xreal_1
a: [3,370,0,2,1,232]
s: 370
i: [1,235,0,3,1,1,0,0] t34_xreal_1
a: [3,355,0,3,1,235]
s: 355
i: [1,234,0,3,1,1,0,0] t34_xreal_1
a: [3,360,0,3,1,234]
s: 360
i: [1,236,0,5,1,1,0,0] t34_xreal_1
a: [3,350,0,5,1,236]
s: 350
i: [1,237,0,2,1,0,0,255] fc1_numbers
a: [3,375,0,2,1,237]
s: 375
i: [1,238,0,2,1,0,0,2] fc1_xboole_0
a: [3,380,0,2,1,238]
s: 380
i: [1,245,0,2,1,0,0,3] fc6_ordinal1
a: [3,410,0,2,1,245]
s: 410
i: [1,244,0,2,1,0,0,3] fc6_ordinal1
a: [3,415,0,2,1,244]
s: 415
i: [1,314,0,2,1,0,0,2] rc2_xboole_0
a: [3,427,0,2,1,314]
s: 427
i: [1,315,0,2,1,0,0,4] rc4_ordinal1
a: [3,432,0,2,1,315]
s: 432
i: [1,316,0,2,1,0,0,4] rc1_ordinal1
a: [3,437,0,2,1,316]
s: 437
i: [1,317,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,442,0,2,1,317]
s: 442
i: [1,318,0,2,1,0,0,1] rc1_xreal_0
a: [3,447,0,2,1,318]
s: 447
i: [1,319,0,2,1,0,0,2] rc1_xxreal_0
a: [3,452,0,2,1,319]
s: 452
i: [1,320,0,2,1,0,0,2] rc1_xboole_0
a: [3,457,0,2,1,320]
s: 457
i: [1,324,0,2,1,0,0,255] rc2_ordinal1
a: [3,462,0,2,1,324]
s: 462
i: [1,323,0,2,1,0,0,255] rc2_ordinal1
a: [3,467,0,2,1,323]
s: 467
i: [1,322,0,2,1,0,0,255] rc2_ordinal1
a: [3,472,0,2,1,322]
s: 472
i: [1,321,0,2,1,0,0,255] rc2_ordinal1
a: [3,477,0,2,1,321]
s: 477
i: [1,326,0,2,1,0,0,4] rc5_ordinal1
a: [3,482,0,2,1,326]
s: 482
i: [1,325,0,2,1,0,0,4] rc5_ordinal1
a: [3,487,0,2,1,325]
s: 487
i: [1,328,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,492,0,2,1,328]
s: 492
i: [1,327,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,497,0,2,1,327]
s: 497
i: [1,332,0,2,1,0,0,2] rc3_xreal_0
a: [3,502,0,2,1,332]
s: 502
i: [1,331,0,2,1,0,0,2] rc3_xreal_0
a: [3,507,0,2,1,331]
s: 507
i: [1,330,0,2,1,0,0,2] rc3_xreal_0
a: [3,512,0,2,1,330]
s: 512
i: [1,329,0,2,1,0,0,2] rc3_xreal_0
a: [3,517,0,2,1,329]
s: 517
i: [1,336,0,2,1,0,0,2] rc2_xreal_0
a: [3,522,0,2,1,336]
s: 522
i: [1,335,0,2,1,0,0,2] rc2_xreal_0
a: [3,527,0,2,1,335]
s: 527
i: [1,334,0,2,1,0,0,2] rc2_xreal_0
a: [3,532,0,2,1,334]
s: 532
i: [1,333,0,2,1,0,0,2] rc2_xreal_0
a: [3,537,0,2,1,333]
s: 537
i: [1,338,0,2,1,0,0,2] rc2_xxreal_0
a: [3,542,0,2,1,338]
s: 542
i: [1,337,0,2,1,0,0,2] rc2_xxreal_0
a: [3,547,0,2,1,337]
s: 547
i: [1,340,0,2,1,0,0,2] rc3_xxreal_0
a: [3,552,0,2,1,340]
s: 552
i: [1,339,0,2,1,0,0,2] rc3_xxreal_0
a: [3,557,0,2,1,339]
s: 557
i: [1,344,0,2,1,0,0,2] rc4_xreal_0
a: [3,562,0,2,1,344]
s: 562
i: [1,343,0,2,1,0,0,2] rc4_xreal_0
a: [3,567,0,2,1,343]
s: 567
i: [1,342,0,2,1,0,0,2] rc4_xreal_0
a: [3,572,0,2,1,342]
s: 572
i: [1,341,0,2,1,0,0,2] rc4_xreal_0
a: [3,577,0,2,1,341]
s: 577
i: [1,346,0,2,1,0,0,2] rc4_xxreal_0
a: [3,582,0,2,1,346]
s: 582
i: [1,345,0,2,1,0,0,2] rc4_xxreal_0
a: [3,587,0,2,1,345]
s: 587
i: [1,296,0,3,1,0,0,255] reflexivity_r1_tarski
s: 296
i: [1,239,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,385,0,3,1,239]
s: 385
i: [1,240,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,390,0,3,1,240]
s: 390
r: 380
d: [2,589,0,2,1,2,23,380,390]
a: [3,593,0,2,1,589]
s: 593
i: [1,241,0,3,1,0,0,2] t1_numerals
a: [3,395,0,3,1,241]
s: 395
d: [2,595,0,3,1,2,22,395,390]
a: [3,599,0,3,1,595]
s: 599
i: [1,247,0,4,2,0,0,4] cc6_ordinal1
s: 247
i: [1,248,0,4,2,0,0,4] cc2_xcmplx_0
s: 248
i: [1,249,0,4,2,0,0,4] cc2_xreal_0
s: 249
d: [2,601,1,2,1,1,40,247,432]
a: [3,606,1,2,1,601]
s: 606
i: [1,250,0,4,2,0,0,4] cc2_xxreal_0
s: 250
d: [2,602,1,2,1,1,40,247,482]
a: [3,611,1,2,1,602]
s: 611
i: [1,251,0,4,2,0,0,255] cc1_ordinal1
s: 251
i: [1,252,0,4,2,0,0,255] cc1_ordinal1
s: 252
d: [2,613,1,2,1,1,40,248,432]
a: [3,618,1,2,1,613]
s: 618
i: [1,255,0,4,2,0,0,2] cc3_xreal_0
s: 255
d: [2,699,1,2,1,1,40,255,370]
a: [3,708,1,2,1,699]
s: 708
i: [1,256,0,4,2,0,0,2] cc4_xreal_0
s: 256
i: [1,265,0,4,2,0,0,255] cc9_ordinal1
s: 265
d: [2,700,1,2,1,1,40,255,365]
a: [3,713,1,2,1,700]
s: 713
i: [1,266,0,4,2,0,0,4] cc7_ordinal1
s: 266
d: [2,720,1,2,1,1,40,256,370]
a: [3,729,1,2,1,720]
s: 729
i: [1,267,0,4,2,0,0,4] cc3_ordinal1
s: 267
i: [1,295,0,4,1,0,0,3] existence_m1_subset_1
s: 295
d: [2,721,1,2,1,1,40,256,365]
a: [3,734,1,2,1,721]
s: 734
i: [1,242,0,4,1,0,0,255] dt_k5_numbers
a: [3,400,0,4,1,242]
s: 400
d: [2,614,1,2,1,1,40,248,482]
a: [3,623,1,2,1,614]
s: 623
i: [1,243,0,4,1,0,0,255] dt_k6_numbers
a: [3,405,0,4,1,243]
s: 405
d: [2,813,0,4,1,2,22,400,385]
a: [3,817,0,4,1,813]
s: 817
d: [2,625,1,2,1,1,40,249,432]
a: [3,630,1,2,1,625]
s: 630
d: [2,819,0,4,1,2,22,405,385]
a: [3,823,0,4,1,819]
s: 823
d: [2,626,1,2,1,1,40,249,482]
a: [3,635,1,2,1,626]
s: 635
i: [1,268,0,5,2,0,0,2] t6_boole
s: 268
i: [1,269,0,5,2,0,0,3] cc8_ordinal1
s: 269
d: [2,637,1,2,1,1,40,250,432]
a: [3,642,1,2,1,637]
s: 642
i: [1,270,0,5,2,0,0,255] cc1_xcmplx_0
s: 270
d: [2,638,1,2,1,1,40,250,482]
a: [3,647,1,2,1,638]
s: 647
i: [1,271,0,5,2,0,0,255] cc1_xreal_0
s: 271
i: [1,292,0,5,2,0,0,255] t7_boole
s: 292
d: [2,649,1,2,1,1,40,251,410]
a: [3,657,1,2,1,649]
s: 657
i: [1,307,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,422,0,5,2,307]
s: 422
d: [2,651,1,2,1,1,40,251,437]
a: [3,667,1,2,1,651]
s: 667
d: [2,831,0,5,2,1,22,268,390]
s: 831
i: [1,276,0,6,3,0,0,255] cc2_ordinal1
s: 276
d: [2,854,1,5,2,2,44,355,831]
s: 854
i: [1,279,0,6,3,0,0,2] cc3_xxreal_0
s: 279
d: [2,860,2,4,2,3,40,854,422]
d: [2,861,2,2,1,4,21,860,734]
a: [3,865,2,2,1,861]
s: 865
i: [1,281,0,6,3,0,0,2] cc3_xxreal_0
s: 281
i: [1,282,0,6,3,0,0,2] cc5_xxreal_0
s: 282
d: [2,855,1,5,2,2,44,360,831]
s: 855
i: [1,284,0,6,3,0,0,2] cc5_xxreal_0
s: 284
d: [2,853,1,7,2,2,44,350,831]
s: 853
i: [1,286,0,6,3,0,0,2] cc7_xxreal_0
s: 286
i: [1,287,0,6,3,0,0,2] cc7_xxreal_0
s: 287
d: [2,674,1,2,1,1,40,252,410]
a: [3,682,1,2,1,674]
s: 682
i: [1,288,0,6,2,0,0,255] t1_subset
s: 288
d: [2,676,1,2,1,1,40,252,437]
a: [3,692,1,2,1,676]
s: 692
i: [1,297,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 297
i: [1,253,0,7,3,0,0,4] cc5_ordinal1
s: 253
d: [2,701,1,2,1,1,40,255,447]
a: [3,718,1,2,1,701]
s: 718
i: [1,254,0,7,2,0,0,1] t1_arithm
s: 254
d: [2,722,1,2,1,1,40,256,447]
a: [3,739,1,2,1,722]
s: 739
i: [1,290,0,7,2,0,0,255] t3_subset
s: 290
i: [1,291,0,7,3,0,0,2] t8_boole
s: 291
d: [2,741,1,2,1,1,40,265,593]
a: [3,748,1,2,1,741]
s: 748
i: [1,246,0,8,3,0,0,255] cc10_ordinal1
s: 246
d: [2,742,1,2,1,1,40,265,457]
a: [3,753,1,2,1,742]
s: 753
i: [1,273,0,8,4,0,0,2] cc4_xxreal_0
s: 273
i: [1,275,0,8,4,0,0,2] cc6_xxreal_0
s: 275
d: [2,932,2,6,3,1,40,273,729]
a: [3,947,2,2,1,932]
s: 947
r: 947
i: [1,277,0,8,4,0,0,2] cc8_xxreal_0
s: 277
a: [3,946,2,2,1,932]
s: 946
i: [1,289,0,8,3,0,0,255] t2_subset
s: 289
i: [1,304,0,8,3,0,0,1] fc2_xcmplx_0
s: 304
a: [3,951,2,2,1,932]
s: 951
i: [1,306,0,8,3,0,0,1] fc5_xreal_0
s: 306
d: [2,933,2,6,3,1,40,273,734]
d: [2,957,2,4,2,2,21,933,865]
a: [3,961,2,2,1,957]
s: 961
i: [1,294,0,9,3,0,0,255] t5_subset
s: 294
i: [1,293,0,10,3,0,0,255] t4_subset
s: 293
d: [2,743,1,2,1,1,40,265,577]
a: [3,758,1,2,1,743]
s: 758
i: [1,308,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 308
d: [2,1066,1,7,3,1,40,308,355]
d: [2,1070,1,5,2,2,21,1066,734]
a: [3,1074,1,2,1,1070]
s: 1074
i: [1,259,0,11,5,0,0,1] t2_real
s: 259
i: [1,260,0,11,5,0,0,1] t1_real
s: 260
a: [3,954,2,2,1,932]
s: 954
i: [1,261,0,11,5,0,0,1] t7_real
s: 261
d: [2,1083,1,12,4,1,40,259,350]
a: [3,1102,1,4,1,1083]
s: 1102
d: [2,1147,2,0,0,3,52,370,365,1102,306]
e: 1147
d: [2,1148,2,4,2,1,40,1102,306]
d: [2,1150,2,2,1,2,21,1148,365]
d: [2,1151,2,0,0,3,21,1150,370]
e: 1151
r: 1102
i: [1,262,0,11,5,0,0,1] t8_real
s: 262
i: [1,263,0,11,5,0,0,1] t4_real
s: 263
d: [2,1082,1,8,4,1,40,259,360]
d: [2,1089,1,6,3,2,21,1082,370]
a: [3,1093,1,2,1,1089]
s: 1093
i: [1,264,0,11,5,0,0,1] t3_real
s: 264
a: [3,1101,1,4,1,1083]
s: 1101
i: [1,305,0,11,3,0,0,2] commutativity_k2_xcmplx_0
s: 305
i: [1,302,0,12,5,0,0,1] fc10_xreal_0
s: 302
a: [3,1106,1,4,1,1083]
s: 1106
i: [1,303,0,12,5,0,0,1] fc9_xreal_0
s: 303
a: [3,965,2,2,1,957]
s: 965
d: [2,1210,3,0,0,3,52,734,961,965,281]
e: 1210
d: [2,1212,3,4,2,1,40,965,281]
d: [2,1214,3,2,1,2,21,1212,961]
d: [2,1215,3,0,0,3,21,1214,734]
e: 1215
r: 965
r: 1106
i: [1,309,0,12,5,0,0,1] fc11_xreal_0
s: 309
i: [1,310,0,12,5,0,0,1] fc14_xreal_0
s: 310
a: [3,1105,1,4,1,1083]
s: 1105
i: [1,311,0,12,5,0,0,1] fc13_xreal_0
s: 311
d: [2,1120,1,12,4,1,40,260,350]
a: [3,1133,1,4,1,1120]
s: 1133
r: 946
r: 951
r: 1133
i: [1,312,0,12,5,0,0,1] fc12_xreal_0
s: 312
i: [1,257,0,13,6,0,0,1] t6_real
s: 257
s: 947
i: [1,258,0,13,6,0,0,1] t5_real
s: 258
s: 951
d: [2,1302,3,0,0,3,52,729,947,951,279]
e: 1302
d: [2,1303,3,4,2,1,40,951,279]
d: [2,1305,3,2,1,2,21,1303,947]
d: [2,1306,3,0,0,3,21,1305,729]
e: 1306
r: 947
i: [1,301,0,14,4,0,0,255] existence_m2_subset_1
s: 301
i: [1,298,0,15,5,0,0,255] dt_m2_subset_1
s: 298
s: 946
i: [1,299,0,15,5,0,0,255] redefinition_m2_subset_1
s: 299
d: [2,1016,3,2,1,1,40,947,267]
a: [3,1022,3,2,1,1016]
s: 1022
r: 1022
i: [1,300,0,15,5,0,0,255] redefinition_m2_subset_1
s: 300
i: [1,313,0,17,4,0,0,2] spc6_arithm
s: 313
a: [3,1021,3,2,1,1016]
s: 1021
a: [3,1132,1,4,1,1120]
s: 1132
d: [2,1017,3,2,1,1,40,947,266]
a: [3,1027,3,2,1,1017]
s: 1027
d: [2,1395,4,0,0,2,52,1021,1027,247]
e: 1395
d: [2,1399,4,2,1,1,40,1027,247]
d: [2,1401,4,0,0,2,21,1399,1021]
e: 1401
r: 1027
d: [2,744,1,2,1,1,40,265,587]
a: [3,763,1,2,1,744]
s: 763
d: [2,765,1,2,1,1,40,266,593]
a: [3,772,1,2,1,765]
s: 772
d: [2,1403,2,0,0,2,52,1074,772,250]
e: 1403
d: [2,1404,2,0,0,2,52,1093,772,249]
e: 1404
d: [2,1405,2,2,1,1,40,772,250]
d: [2,1412,2,0,0,2,21,1405,1074]
e: 1412
d: [2,1406,2,2,1,1,40,772,249]
d: [2,1414,2,0,0,2,21,1406,1093]
e: 1414
r: 1074
r: 1093
r: 961
r: 1105
r: 1132
a: [3,1073,1,2,1,1070]
s: 1073
a: [3,1092,1,2,1,1089]
s: 1092
a: [3,1096,1,2,1,1089]
s: 1096
a: [3,1129,1,2,1,1120]
s: 1129
d: [2,1442,2,0,0,3,52,1073,593,1129,279]
e: 1442
d: [2,1443,2,4,2,1,40,1129,279]
d: [2,1446,2,2,1,2,21,1443,593]
d: [2,1447,2,0,0,3,21,1446,1073]
e: 1447
r: 1096
r: 1129
a: [3,1097,1,2,1,1089]
s: 1097
d: [2,1449,2,0,0,3,52,1073,593,1097,282]
e: 1449
d: [2,1450,2,4,2,1,40,1097,282]
d: [2,1454,2,2,1,2,21,1450,593]
d: [2,1455,2,0,0,3,21,1454,1073]
e: 1455
r: 946
r: 951
r: 1021
r: 1097
s: 1096
a: [3,1128,1,2,1,1120]
s: 1128
s: 947
a: [3,1078,1,3,1,1070]
s: 1078
s: 1105
a: [3,950,2,2,1,932]
s: 950
s: 1133
d: [2,832,1,2,1,1,40,269,599]
d: [2,1141,1,8,4,1,40,261,355]
d: [2,1429,1,6,3,2,21,1141,365]
d: [2,1493,2,8,4,1,40,1133,312]
d: [2,1495,2,6,3,2,21,1493,370]
d: [2,1496,2,4,2,3,21,1495,954]
d: [2,1497,2,2,1,4,21,1496,365]
f: 1497,1455,1447,1429,1406,1405,1151,1120,1089,832,595,589,241,240,238,232,233,234,235,236
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t34_xreal_1
% SZS output start Proof for xreal_1__t34_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k2_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0)))) [input t34_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k2_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0)))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
12. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
15. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
21. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & v2_xxreal_0(X0)) => v2_xxreal_0(k2_xcmplx_0(X1,X0))) [input fc12_xreal_0]
32. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
34. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
41. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v3_xxreal_0(X1) & ~v2_xxreal_0(X0) & ~r1_xxreal_0(X0,X1)))) [input t7_real]
53. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
61. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
74. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
85. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
90. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k2_xcmplx_0(X0,X1))) [input fc5_xreal_0]
103. ? [X0] : (? [X1] : ((r1_xxreal_0(k2_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
104. ? [X0] : (? [X1] : (r1_xxreal_0(k2_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 103]
108. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
109. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
119. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 34]
120. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 119]
121. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 32]
122. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 121]
123. ! [X0] : (! [X1] : ((v3_xxreal_0(X1) | v2_xxreal_0(X0) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 41]
124. ! [X0] : (! [X1] : (v3_xxreal_0(X1) | v2_xxreal_0(X0) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 123]
135. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 53]
146. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 12]
147. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 146]
148. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 15]
149. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 148]
176. ! [X0,X1] : (v1_xreal_0(k2_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 90]
177. ! [X0,X1] : (v1_xreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 176]
188. ! [X0,X1] : (v2_xxreal_0(k2_xcmplx_0(X1,X0)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0))) [ennf transformation 21]
189. ! [X0,X1] : (v2_xxreal_0(k2_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0)) [flattening 188]
192. ? [X0] : (? [X1] : (r1_xxreal_0(k2_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (r1_xxreal_0(k2_xcmplx_0(sK0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
193. ? [X1] : (r1_xxreal_0(k2_xcmplx_0(sK0,X1),k6_numbers) & ~r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) => (r1_xxreal_0(k2_xcmplx_0(sK0,sK1),k6_numbers) & ~r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) [choice axiom]
194. (r1_xxreal_0(k2_xcmplx_0(sK0,sK1),k6_numbers) & ~r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 104,193,192]
232. v1_xreal_0(sK0) [cnf transformation 194] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
233. v1_xreal_0(sK1) [cnf transformation 194] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
234. r1_xxreal_0(k6_numbers,sK0) [cnf transformation 194] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
235. ~r1_xxreal_0(sK1,k6_numbers) [cnf transformation 194] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
236. r1_xxreal_0(k2_xcmplx_0(sK0,sK1),k6_numbers) [cnf transformation 194] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
238. v1_xboole_0(k1_xboole_0) [cnf transformation 74] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
240. k6_numbers = k1_xboole_0 [cnf transformation 85] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
241. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 61] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
249. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 108] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
250. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 109] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
259. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 120] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
260. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 122] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
261. r1_xxreal_0(X0,X1) | v2_xxreal_0(X0) | v3_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 124] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
269. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 135] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
279. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 147] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
282. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 149] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
306. v1_xreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 177] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
312. v2_xxreal_0(k2_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0) [cnf transformation 189] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
348. 1 <=> r1_xxreal_0(k2_xcmplx_0(sK0,sK1),k6_numbers) [avatar definition]
350. r1_xxreal_0(k2_xcmplx_0(sK0,sK1),k6_numbers) <- (1) [avatar component clause 348] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
351. 1 [avatar split clause 236,348]
353. 2 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
355. ~r1_xxreal_0(sK1,k6_numbers) <- (~2) [avatar component clause 353] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
356. ~2 [avatar split clause 235,353]
358. 3 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
360. r1_xxreal_0(k6_numbers,sK0) <- (3) [avatar component clause 358] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
361. 3 [avatar split clause 234,358]
363. 4 <=> v1_xreal_0(sK1) [avatar definition]
365. v1_xreal_0(sK1) <- (4) [avatar component clause 363] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
366. 4 [avatar split clause 233,363]
368. 5 <=> v1_xreal_0(sK0) [avatar definition]
370. v1_xreal_0(sK0) <- (5) [avatar component clause 368] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
371. 5 [avatar split clause 232,368]
378. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
380. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 378] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
381. 7 [avatar split clause 238,378]
388. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
390. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 388] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
391. 9 [avatar split clause 240,388]
393. 10 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
395. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (10) [avatar component clause 393] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
396. 10 [avatar split clause 241,393]
589. v1_xboole_0(k6_numbers) <- (7, 9) [backward demodulation 380,390] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
591. 50 <=> v1_xboole_0(k6_numbers) [avatar definition]
593. v1_xboole_0(k6_numbers) <- (50) [avatar component clause 591] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
594. 50 | ~7 | ~9 [avatar split clause 589,388,378,591]
595. m1_subset_1(k6_numbers,k4_ordinal1) <- (9, 10) [forward demodulation 395,390] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
597. 51 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
599. m1_subset_1(k6_numbers,k4_ordinal1) <- (51) [avatar component clause 597] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
600. 51 | ~9 | ~10 [avatar split clause 595,393,388,597]
770. 78 <=> v7_ordinal1(k6_numbers) [avatar definition]
772. v7_ordinal1(k6_numbers) <- (78) [avatar component clause 770] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
832. v7_ordinal1(k6_numbers) <- (51) [resolution 269,599] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
834. 78 | ~51 [avatar split clause 832,597,770]
953. 102 <=> v3_xxreal_0(sK0) [avatar definition]
954. ~v3_xxreal_0(sK0) <- (~102) [avatar component clause 953] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
959. 103 <=> v2_xxreal_0(sK1) [avatar definition]
1072. 122 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1073. v1_xxreal_0(k6_numbers) <- (122) [avatar component clause 1072] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1082. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (3) [resolution 259,360] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1089. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 5) [subsumption resolution 1082,370] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1091. 124 <=> v1_xreal_0(k6_numbers) [avatar definition]
1095. 125 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1097. v3_xxreal_0(k6_numbers) <- (125) [avatar component clause 1095] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1098. ~124 | 125 | ~102 | ~3 | ~5 [avatar split clause 1089,368,358,953,1095,1091]
1100. 126 <=> v1_xreal_0(k2_xcmplx_0(sK0,sK1)) [avatar definition]
1102. ~v1_xreal_0(k2_xcmplx_0(sK0,sK1)) <- (~126) [avatar component clause 1100] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1120. ~v2_xxreal_0(k2_xcmplx_0(sK0,sK1)) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k2_xcmplx_0(sK0,sK1)) <- (1) [resolution 260,350] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1127. 130 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1129. v2_xxreal_0(k6_numbers) <- (130) [avatar component clause 1127] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1131. 131 <=> v2_xxreal_0(k2_xcmplx_0(sK0,sK1)) [avatar definition]
1133. ~v2_xxreal_0(k2_xcmplx_0(sK0,sK1)) <- (~131) [avatar component clause 1131] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1134. ~126 | ~124 | 130 | ~131 | ~1 [avatar split clause 1120,348,1131,1127,1091,1100]
1141. v2_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (~2) [resolution 261,355] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1148. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~126) [resolution 1102,306] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1150. ~v1_xreal_0(sK0) <- (4, ~126) [subsumption resolution 1148,365] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1151. $false <- (4, 5, ~126) [subsumption resolution 1150,370] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1152. ~4 | ~5 | 126 [avatar contradiction clause 1151]
1405. v1_xxreal_0(k6_numbers) <- (78) [resolution 772,250] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1406. v1_xreal_0(k6_numbers) <- (78) [resolution 772,249] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1416. 122 | ~78 [avatar split clause 1405,770,1072]
1417. 124 | ~78 [avatar split clause 1406,770,1091]
1429. v2_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (~2, 4) [subsumption resolution 1141,365] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1430. ~124 | 125 | 103 | 2 | ~4 [avatar split clause 1429,363,353,959,1095,1091]
1443. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (130) [resolution 1129,279] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1446. ~v1_xxreal_0(k6_numbers) <- (50, 130) [subsumption resolution 1443,593] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1447. $false <- (50, 122, 130) [subsumption resolution 1446,1073] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1448. ~50 | ~122 | ~130 [avatar contradiction clause 1447]
1450. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (125) [resolution 1097,282] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1454. ~v1_xxreal_0(k6_numbers) <- (50, 125) [subsumption resolution 1450,593] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1455. $false <- (50, 122, 125) [subsumption resolution 1454,1073] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1456. ~50 | ~122 | ~125 [avatar contradiction clause 1455]
1493. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) | ~v1_xreal_0(sK1) | ~v2_xxreal_0(sK1) <- (~131) [resolution 1133,312] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1495. v3_xxreal_0(sK0) | ~v1_xreal_0(sK1) | ~v2_xxreal_0(sK1) <- (5, ~131) [subsumption resolution 1493,370] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1496. ~v1_xreal_0(sK1) | ~v2_xxreal_0(sK1) <- (5, ~102, ~131) [subsumption resolution 1495,954] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1497. ~v2_xxreal_0(sK1) <- (4, 5, ~102, ~131) [subsumption resolution 1496,365] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1498. ~103 | ~4 | ~5 | 102 | 131 [avatar split clause 1497,1131,953,368,363,959]
1499. $false [avatar sat refutation 351,356,361,366,371,381,391,396,594,600,834,1098,1134,1152,1416,1417,1430,1448,1456,1498] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1064,thDist:-1064}
% SZS output end Proof for xreal_1__t34_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 90
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 115
% Generated clauses: 705
% Active clauses: 161
% Passive clauses: 292
% Final active clauses: 139
% Final passive clauses: 103
% Main loop iterations started: 173
% 
% >>> Simplifying Inferences
% Duplicate literals: 47
% Fw subsumption resolutions: 98
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 23
% Forward subsumptions: 80
% 
% >>> Generating Inferences
% Binary resolution: 190
% Unit resulting resolution: 8
% Factoring: 1
% Forward superposition: 41
% Backward superposition: 23
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 43
% Split components: 106
% Sat splitting refutations: 16
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 186
% SAT solver unit clauses: 50
% SAT solver binary clauses: 71
% 
% Memory used [KB]: 5756
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s ( own 0.0000 s ) 
% term sharing: 0.002 s
% splitting model update: 0.001 s ( own 0.0000 s ) 
% unit clause index maintenance: 0.001 s
% backward subsumption index maintenance: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.007 s ( own 0.002 s ) 

