i: [1,24,0,3,1,1,0,0] l3_jordan
a: [3,41,0,3,1,24]
s: 41
i: [1,23,0,3,1,1,0,0] l3_jordan
a: [3,46,0,3,1,23]
s: 46
i: [1,22,0,3,1,1,0,0] l3_jordan
a: [3,51,0,3,1,22]
s: 51
i: [1,21,0,3,1,1,0,0] l3_jordan
a: [3,56,0,3,1,21]
s: 56
i: [1,25,0,9,1,1,0,0] l3_jordan
a: [3,36,0,9,1,25]
s: 36
i: [1,30,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 30
i: [1,32,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 32
i: [1,31,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 31
r: 36
d: [2,58,1,3,1,1,40,32,56]
a: [3,65,1,3,1,58]
s: 65
d: [2,82,0,9,1,1,23,36,31]
d: [2,83,0,9,1,1,22,82,31]
a: [3,87,0,9,1,83]
s: 87
i: [1,27,0,8,2,0,0,1] t70_xboole_1
s: 27
d: [2,59,1,3,1,1,40,32,51]
a: [3,70,1,3,1,59]
s: 70
i: [1,28,0,8,2,0,0,1] t70_xboole_1
s: 28
d: [2,60,1,3,1,1,40,32,46]
a: [3,75,1,3,1,60]
s: 75
i: [1,26,0,11,3,0,0,1] t70_xboole_1
s: 26
i: [1,29,0,16,4,0,0,1] t114_xboole_1
s: 29
d: [2,61,1,3,1,1,40,32,41]
a: [3,80,1,3,1,61]
s: 80
d: [2,91,1,8,2,0,44,27,31]
s: 91
d: [2,92,1,8,2,0,44,27,31]
s: 92
d: [2,100,1,11,3,0,40,26,32]
s: 100
d: [2,108,1,14,4,0,44,29,30]
d: [2,114,1,11,3,0,18,108]
s: 114
d: [2,116,2,10,2,1,40,100,87]
d: [2,124,2,7,1,2,21,116,80]
a: [3,128,2,7,1,124]
s: 128
d: [2,105,1,16,4,0,40,29,32]
s: 105
d: [2,144,3,8,2,1,40,128,26]
d: [2,145,3,5,1,2,21,144,75]
a: [3,149,3,5,1,145]
s: 149
d: [2,161,4,0,0,3,52,70,65,149,26]
e: 161
d: [2,162,4,6,2,1,40,149,26]
d: [2,164,4,3,1,2,21,162,65]
d: [2,165,4,0,0,3,21,164,70]
e: 165
f: 165,145,124,83,61,60,59,58,21,22,23,24,25
% Refutation found. Thanks to Tanya!
% SZS status Theorem for jordan__l3_jordan
% SZS output start Proof for jordan__l3_jordan
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((r1_xboole_0(X3,X4) & r1_xboole_0(X2,X4) & r1_xboole_0(X1,X4) & r1_xboole_0(X0,X4)) => r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3),X4)) [input l3_jordan]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((r1_xboole_0(X3,X4) & r1_xboole_0(X2,X4) & r1_xboole_0(X1,X4) & r1_xboole_0(X0,X4)) => r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3),X4)) [negated conjecture 1]
3. ! [X0,X1] : k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [input commutativity_k2_xboole_0]
6. ! [X0,X1] : (r1_xboole_0(X0,X1) => r1_xboole_0(X1,X0)) [input symmetry_r1_xboole_0]
8. ! [X0] : ! [X1] : ! [X2] : (~(r1_xboole_0(X0,k2_xboole_0(X1,X2)) & ~(r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1))) & ~(r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1) & ~r1_xboole_0(X0,k2_xboole_0(X1,X2)))) [input t70_xboole_1]
9. ~! [X0,X1,X2,X3,X4] : ((r1_xboole_0(X3,X4) & r1_xboole_0(X2,X4) & r1_xboole_0(X1,X4) & r1_xboole_0(X0,X4)) => r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3),X4)) [flattening 2]
10. ! [X0,X1,X2] : (~(r1_xboole_0(X0,k2_xboole_0(X1,X2)) & ~(r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1))) & ~(r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1) & ~r1_xboole_0(X0,k2_xboole_0(X1,X2)))) [flattening 8]
13. ? [X0,X1,X2,X3,X4] : (~r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3),X4) & (r1_xboole_0(X3,X4) & r1_xboole_0(X2,X4) & r1_xboole_0(X1,X4) & r1_xboole_0(X0,X4))) [ennf transformation 9]
14. ? [X0,X1,X2,X3,X4] : (~r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3),X4) & r1_xboole_0(X3,X4) & r1_xboole_0(X2,X4) & r1_xboole_0(X1,X4) & r1_xboole_0(X0,X4)) [flattening 13]
15. ! [X0,X1,X2] : ((~r1_xboole_0(X0,k2_xboole_0(X1,X2)) | (r1_xboole_0(X0,X2) & r1_xboole_0(X0,X1))) & (~r1_xboole_0(X0,X2) | ~r1_xboole_0(X0,X1) | r1_xboole_0(X0,k2_xboole_0(X1,X2)))) [ennf transformation 10]
18. ! [X0,X1] : (r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 6]
19. ? [X0,X1,X2,X3,X4] : (~r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(X0,X1),X2),X3),X4) & r1_xboole_0(X3,X4) & r1_xboole_0(X2,X4) & r1_xboole_0(X1,X4) & r1_xboole_0(X0,X4)) => (~r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3),sK4) & r1_xboole_0(sK3,sK4) & r1_xboole_0(sK2,sK4) & r1_xboole_0(sK1,sK4) & r1_xboole_0(sK0,sK4)) [choice axiom]
20. ~r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3),sK4) & r1_xboole_0(sK3,sK4) & r1_xboole_0(sK2,sK4) & r1_xboole_0(sK1,sK4) & r1_xboole_0(sK0,sK4) [skolemisation 14,19]
21. r1_xboole_0(sK0,sK4) [cnf transformation 20] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
22. r1_xboole_0(sK1,sK4) [cnf transformation 20] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
23. r1_xboole_0(sK2,sK4) [cnf transformation 20] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
24. r1_xboole_0(sK3,sK4) [cnf transformation 20] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
25. ~r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3),sK4) [cnf transformation 20] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
26. r1_xboole_0(X0,k2_xboole_0(X1,X2)) | ~r1_xboole_0(X0,X1) | ~r1_xboole_0(X0,X2) [cnf transformation 15] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
31. k2_xboole_0(X0,X1) = k2_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
32. ~r1_xboole_0(X0,X1) | r1_xboole_0(X1,X0) [cnf transformation 18] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. 1 <=> r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3),sK4) [avatar definition]
36. ~r1_xboole_0(k2_xboole_0(k2_xboole_0(k2_xboole_0(sK0,sK1),sK2),sK3),sK4) <- (~1) [avatar component clause 34] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
37. ~1 [avatar split clause 25,34]
39. 2 <=> r1_xboole_0(sK3,sK4) [avatar definition]
41. r1_xboole_0(sK3,sK4) <- (2) [avatar component clause 39] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
42. 2 [avatar split clause 24,39]
44. 3 <=> r1_xboole_0(sK2,sK4) [avatar definition]
46. r1_xboole_0(sK2,sK4) <- (3) [avatar component clause 44] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
47. 3 [avatar split clause 23,44]
49. 4 <=> r1_xboole_0(sK1,sK4) [avatar definition]
51. r1_xboole_0(sK1,sK4) <- (4) [avatar component clause 49] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
52. 4 [avatar split clause 22,49]
54. 5 <=> r1_xboole_0(sK0,sK4) [avatar definition]
56. r1_xboole_0(sK0,sK4) <- (5) [avatar component clause 54] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
57. 5 [avatar split clause 21,54]
58. r1_xboole_0(sK4,sK0) <- (5) [resolution 32,56] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
59. r1_xboole_0(sK4,sK1) <- (4) [resolution 32,51] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
60. r1_xboole_0(sK4,sK2) <- (3) [resolution 32,46] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
61. r1_xboole_0(sK4,sK3) <- (2) [resolution 32,41] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
63. 6 <=> r1_xboole_0(sK4,sK0) [avatar definition]
65. r1_xboole_0(sK4,sK0) <- (6) [avatar component clause 63] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
66. 6 | ~5 [avatar split clause 58,54,63]
68. 7 <=> r1_xboole_0(sK4,sK1) [avatar definition]
70. r1_xboole_0(sK4,sK1) <- (7) [avatar component clause 68] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
71. 7 | ~4 [avatar split clause 59,49,68]
73. 8 <=> r1_xboole_0(sK4,sK2) [avatar definition]
75. r1_xboole_0(sK4,sK2) <- (8) [avatar component clause 73] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
76. 8 | ~3 [avatar split clause 60,44,73]
78. 9 <=> r1_xboole_0(sK4,sK3) [avatar definition]
80. r1_xboole_0(sK4,sK3) <- (9) [avatar component clause 78] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
81. 9 | ~2 [avatar split clause 61,39,78]
82. ~r1_xboole_0(k2_xboole_0(sK3,k2_xboole_0(k2_xboole_0(sK0,sK1),sK2)),sK4) <- (~1) [backward demodulation 36,31] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. ~r1_xboole_0(k2_xboole_0(sK3,k2_xboole_0(sK2,k2_xboole_0(sK0,sK1))),sK4) <- (~1) [forward demodulation 82,31] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
85. 10 <=> r1_xboole_0(k2_xboole_0(sK3,k2_xboole_0(sK2,k2_xboole_0(sK0,sK1))),sK4) [avatar definition]
87. ~r1_xboole_0(k2_xboole_0(sK3,k2_xboole_0(sK2,k2_xboole_0(sK0,sK1))),sK4) <- (~10) [avatar component clause 85] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
88. ~10 | 1 [avatar split clause 83,34,85]
100. r1_xboole_0(k2_xboole_0(X7,X8),X6) | ~r1_xboole_0(X6,X8) | ~r1_xboole_0(X6,X7) [resolution 26,32] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
116. ~r1_xboole_0(sK4,k2_xboole_0(sK2,k2_xboole_0(sK0,sK1))) | ~r1_xboole_0(sK4,sK3) <- (~10) [resolution 100,87] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
124. ~r1_xboole_0(sK4,k2_xboole_0(sK2,k2_xboole_0(sK0,sK1))) <- (9, ~10) [subsumption resolution 116,80] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
126. 11 <=> r1_xboole_0(sK4,k2_xboole_0(sK2,k2_xboole_0(sK0,sK1))) [avatar definition]
128. ~r1_xboole_0(sK4,k2_xboole_0(sK2,k2_xboole_0(sK0,sK1))) <- (~11) [avatar component clause 126] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
129. ~11 | ~9 | 10 [avatar split clause 124,85,78,126]
144. ~r1_xboole_0(sK4,sK2) | ~r1_xboole_0(sK4,k2_xboole_0(sK0,sK1)) <- (~11) [resolution 128,26] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
145. ~r1_xboole_0(sK4,k2_xboole_0(sK0,sK1)) <- (8, ~11) [subsumption resolution 144,75] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
147. 13 <=> r1_xboole_0(sK4,k2_xboole_0(sK0,sK1)) [avatar definition]
149. ~r1_xboole_0(sK4,k2_xboole_0(sK0,sK1)) <- (~13) [avatar component clause 147] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
150. ~13 | ~8 | 11 [avatar split clause 145,126,73,147]
162. ~r1_xboole_0(sK4,sK0) | ~r1_xboole_0(sK4,sK1) <- (~13) [resolution 149,26] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
164. ~r1_xboole_0(sK4,sK1) <- (6, ~13) [subsumption resolution 162,65] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
165. $false <- (6, 7, ~13) [subsumption resolution 164,70] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
166. ~6 | ~7 | 13 [avatar contradiction clause 165]
167. $false [avatar sat refutation 37,42,47,52,57,66,71,76,81,88,129,150,166] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:79,thDist:-79}
% SZS output end Proof for jordan__l3_jordan
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 7
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 92
% Active clauses: 22
% Passive clauses: 38
% Final active clauses: 21
% Final passive clauses: 14
% Main loop iterations started: 24
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Fw subsumption resolutions: 5
% Fw demodulations: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 27
% Unit resulting resolution: 1
% Forward superposition: 27
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 5
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% property evaluation: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

