i: [1,37,0,8,2,1,0,0] t11_gate_1
a: [3,50,0,4,1,37]
s: 50
a: [3,55,0,4,1,37]
s: 55
i: [1,46,0,2,1,0,0,1] rc2_xboole_0
a: [3,61,0,2,1,46]
s: 61
i: [1,47,0,2,1,0,0,1] rc1_xboole_0
a: [3,66,0,2,1,47]
s: 66
i: [1,38,0,7,1,0,0,255] commutativity_k4_gate_1
s: 38
i: [1,39,0,7,1,0,0,255] commutativity_k5_gate_1
s: 39
i: [1,40,0,8,3,0,0,1] fc12_gate_1
s: 40
i: [1,41,0,8,3,0,0,1] fc9_gate_1
s: 41
d: [2,70,1,4,2,1,40,41,55]
a: [3,76,1,2,1,70]
s: 76
i: [1,42,0,8,3,0,0,1] fc7_gate_1
s: 42
i: [1,43,0,8,3,0,0,1] fc10_gate_1
s: 43
a: [3,80,1,2,1,70]
s: 80
i: [1,44,0,8,3,0,0,1] fc11_gate_1
s: 44
i: [1,45,0,8,3,0,0,1] fc8_gate_1
s: 45
d: [2,97,2,0,0,3,52,76,80,55,45]
e: 97
d: [2,98,1,4,2,1,40,45,55]
d: [2,104,1,2,1,2,21,98,80]
d: [2,105,1,0,0,3,21,104,76]
e: 105
r: 50
r: 55
a: [3,51,0,4,1,37]
s: 51
d: [2,107,2,0,0,3,52,76,80,51,44]
e: 107
d: [2,108,1,4,2,1,40,51,44]
d: [2,111,1,2,1,2,21,108,80]
d: [2,112,1,0,0,3,21,111,76]
e: 112
r: 76
a: [3,54,0,4,1,37]
s: 54
a: [3,75,1,2,1,70]
s: 75
d: [2,83,1,8,3,0,44,42,38]
s: 83
d: [2,116,2,0,0,3,52,80,75,54,83]
e: 116
d: [2,119,2,4,2,1,40,83,54]
d: [2,125,2,2,1,2,21,119,75]
d: [2,126,2,0,0,3,21,125,80]
e: 126
r: 51
r: 54
s: 50
s: 55
d: [2,84,1,8,3,0,44,42,38]
s: 84
d: [2,88,1,8,3,0,44,43,39]
s: 88
d: [2,131,2,0,0,3,52,80,75,50,88]
e: 131
d: [2,134,2,4,2,1,40,88,50]
d: [2,140,2,2,1,2,21,134,75]
d: [2,141,2,0,0,3,21,140,80]
e: 141
r: 50
r: 55
r: 75
r: 80
s: 51
s: 76
s: 54
i: [1,36,0,8,2,1,0,0] t11_gate_1
d: [2,86,1,4,2,1,40,43,50]
d: [2,91,1,2,1,2,21,86,76]
d: [2,143,1,4,2,1,40,51,44]
d: [2,144,1,4,2,1,40,51,40]
d: [2,148,1,4,2,1,40,54,42]
d: [2,149,1,2,1,2,21,148,76]
f: 149,144,143,141,126,105,91,70,36,37
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t11_gate_1
% SZS output start Proof for gate_1__t11_gate_1
1. ! [X0] : ! [X1] : (~v1_xboole_0(k5_gate_1(X0,X1)) <=> v1_xboole_0(k4_gate_1(X0,X1))) [input t11_gate_1]
2. ~! [X0] : ! [X1] : (~v1_xboole_0(k5_gate_1(X0,X1)) <=> v1_xboole_0(k4_gate_1(X0,X1))) [negated conjecture 1]
3. ! [X0,X1] : k4_gate_1(X0,X1) = k4_gate_1(X1,X0) [input commutativity_k4_gate_1]
4. ! [X0,X1] : k5_gate_1(X0,X1) = k5_gate_1(X1,X0) [input commutativity_k5_gate_1]
7. ! [X0,X1] : ((~v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k5_gate_1(X0,X1))) [input fc10_gate_1]
8. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => ~v1_xboole_0(k5_gate_1(X0,X1))) [input fc11_gate_1]
9. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k5_gate_1(X0,X1))) [input fc12_gate_1]
10. ! [X0,X1] : ((~v1_xboole_0(X1) & v1_xboole_0(X0)) => ~v1_xboole_0(k4_gate_1(X0,X1))) [input fc7_gate_1]
11. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k4_gate_1(X0,X1))) [input fc8_gate_1]
12. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => v1_xboole_0(k4_gate_1(X0,X1))) [input fc9_gate_1]
15. ~! [X0,X1] : (~v1_xboole_0(k5_gate_1(X0,X1)) <=> v1_xboole_0(k4_gate_1(X0,X1))) [flattening 2]
16. ? [X0,X1] : (~v1_xboole_0(k5_gate_1(X0,X1)) <~> v1_xboole_0(k4_gate_1(X0,X1))) [ennf transformation 15]
17. ! [X0,X1] : (~v1_xboole_0(k5_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 9]
18. ! [X0,X1] : (~v1_xboole_0(k5_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 17]
19. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 12]
20. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 19]
21. ! [X0,X1] : (~v1_xboole_0(k4_gate_1(X0,X1)) | (v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 10]
22. ! [X0,X1] : (~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 21]
23. ! [X0,X1] : (v1_xboole_0(k5_gate_1(X0,X1)) | (v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 7]
24. ! [X0,X1] : (v1_xboole_0(k5_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 23]
25. ! [X0,X1] : (~v1_xboole_0(k5_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 8]
26. ! [X0,X1] : (~v1_xboole_0(k5_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 25]
27. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 11]
28. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 27]
29. ? [X0,X1] : ((~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(k5_gate_1(X0,X1))) & (v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(k5_gate_1(X0,X1)))) [nnf transformation 16]
30. ? [X0,X1] : ((~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(k5_gate_1(X0,X1))) & (v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(k5_gate_1(X0,X1)))) => ((~v1_xboole_0(k4_gate_1(sK0,sK1)) | v1_xboole_0(k5_gate_1(sK0,sK1))) & (v1_xboole_0(k4_gate_1(sK0,sK1)) | ~v1_xboole_0(k5_gate_1(sK0,sK1)))) [choice axiom]
31. (~v1_xboole_0(k4_gate_1(sK0,sK1)) | v1_xboole_0(k5_gate_1(sK0,sK1))) & (v1_xboole_0(k4_gate_1(sK0,sK1)) | ~v1_xboole_0(k5_gate_1(sK0,sK1))) [skolemisation 29,30]
36. v1_xboole_0(k4_gate_1(sK0,sK1)) | ~v1_xboole_0(k5_gate_1(sK0,sK1)) [cnf transformation 31] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. ~v1_xboole_0(k4_gate_1(sK0,sK1)) | v1_xboole_0(k5_gate_1(sK0,sK1)) [cnf transformation 31] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. k4_gate_1(X0,X1) = k4_gate_1(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. k5_gate_1(X0,X1) = k5_gate_1(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. ~v1_xboole_0(k5_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 18] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
41. v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 20] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. ~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
43. v1_xboole_0(k5_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 24] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. ~v1_xboole_0(k5_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 28] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. 1 <=> v1_xboole_0(k5_gate_1(sK0,sK1)) [avatar definition]
50. ~v1_xboole_0(k5_gate_1(sK0,sK1)) <- (~1) [avatar component clause 49] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
51. v1_xboole_0(k5_gate_1(sK0,sK1)) <- (1) [avatar component clause 49] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
53. 2 <=> v1_xboole_0(k4_gate_1(sK0,sK1)) [avatar definition]
54. v1_xboole_0(k4_gate_1(sK0,sK1)) <- (2) [avatar component clause 53] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
55. ~v1_xboole_0(k4_gate_1(sK0,sK1)) <- (~2) [avatar component clause 53] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. 1 | ~2 [avatar split clause 37,53,49]
57. ~1 | 2 [avatar split clause 36,53,49]
70. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~2) [resolution 41,55] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
74. 5 <=> v1_xboole_0(sK0) [avatar definition]
75. ~v1_xboole_0(sK0) <- (~5) [avatar component clause 74] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
76. v1_xboole_0(sK0) <- (5) [avatar component clause 74] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
78. 6 <=> v1_xboole_0(sK1) [avatar definition]
80. v1_xboole_0(sK1) <- (6) [avatar component clause 78] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
81. 5 | 6 | 2 [avatar split clause 70,53,78,74]
83. ~v1_xboole_0(k4_gate_1(X1,X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [superposition 42,38] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
86. v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~1) [resolution 43,50] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
88. v1_xboole_0(k5_gate_1(X1,X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [superposition 43,39] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
91. v1_xboole_0(sK1) <- (~1, 5) [subsumption resolution 86,76] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
92. 6 | 1 | ~5 [avatar split clause 91,74,49,78]
98. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~2) [resolution 45,55] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
104. ~v1_xboole_0(sK0) <- (~2, 6) [subsumption resolution 98,80] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
105. $false <- (~2, 5, 6) [subsumption resolution 104,76] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
106. 2 | ~5 | ~6 [avatar contradiction clause 105]
119. v1_xboole_0(sK0) | ~v1_xboole_0(sK1) <- (2) [resolution 83,54] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
125. ~v1_xboole_0(sK1) <- (2, ~5) [subsumption resolution 119,75] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
126. $false <- (2, ~5, 6) [subsumption resolution 125,80] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
127. ~2 | 5 | ~6 [avatar contradiction clause 126]
134. v1_xboole_0(sK0) | ~v1_xboole_0(sK1) <- (~1) [resolution 88,50] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
140. ~v1_xboole_0(sK1) <- (~1, ~5) [subsumption resolution 134,75] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
141. $false <- (~1, ~5, 6) [subsumption resolution 140,80] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
142. 1 | 5 | ~6 [avatar contradiction clause 141]
143. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (1) [resolution 51,44] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
144. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1) [resolution 51,40] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
145. ~5 | ~6 | ~1 [avatar split clause 143,49,78,74]
146. 5 | 6 | ~1 [avatar split clause 144,49,78,74]
148. v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (2) [resolution 54,42] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
149. v1_xboole_0(sK1) <- (2, 5) [subsumption resolution 148,76] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
150. 6 | ~2 | ~5 [avatar split clause 149,74,53,78]
151. $false [avatar sat refutation 56,57,81,92,106,127,142,145,146,150] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:102,thDist:-102}
% SZS output end Proof for gate_1__t11_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 89
% Active clauses: 24
% Passive clauses: 27
% Final active clauses: 15
% Final passive clauses: 2
% Main loop iterations started: 25
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Fw subsumption resolutions: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Forward subsumptions: 19
% 
% >>> Generating Inferences
% Binary resolution: 23
% Unit resulting resolution: 4
% Forward superposition: 16
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4861
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

