i: [1,260,0,2,1,1,0,0] t31_square_1
a: [3,424,0,2,1,260]
s: 424
i: [1,259,0,2,1,1,0,0] t31_square_1
a: [3,429,0,2,1,259]
s: 429
i: [1,265,0,12,3,1,0,0] t31_square_1
a: [3,402,0,3,1,265]
s: 402
r: 429
r: 402
s: 429
a: [3,403,0,3,1,265]
s: 403
a: [3,406,0,3,1,265]
s: 406
r: 424
r: 406
s: 424
a: [3,407,0,3,1,265]
s: 407
i: [1,262,0,3,1,1,0,0] t31_square_1
a: [3,414,0,3,1,262]
s: 414
i: [1,261,0,3,1,1,0,0] t31_square_1
a: [3,419,0,3,1,261]
s: 419
a: [3,398,0,6,1,265]
s: 398
i: [1,266,0,2,1,0,0,3] fc1_numbers
a: [3,434,0,2,1,266]
s: 434
i: [1,267,0,2,1,0,0,255] fc1_xboole_0
a: [3,439,0,2,1,267]
s: 439
i: [1,268,0,2,1,0,0,255] spc0_boole
a: [3,444,0,2,1,268]
s: 444
i: [1,278,0,2,1,0,0,255] fc6_ordinal1
a: [3,489,0,2,1,278]
s: 489
i: [1,277,0,2,1,0,0,255] fc6_ordinal1
a: [3,494,0,2,1,277]
s: 494
i: [1,351,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,513,0,2,1,351]
s: 513
i: [1,358,0,2,1,0,0,2] rc2_xboole_0
a: [3,521,0,2,1,358]
s: 521
i: [1,359,0,2,1,0,0,255] rc4_ordinal1
a: [3,526,0,2,1,359]
s: 526
i: [1,360,0,2,1,0,0,255] rc1_ordinal1
a: [3,531,0,2,1,360]
s: 531
i: [1,361,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,536,0,2,1,361]
s: 536
i: [1,362,0,2,1,0,0,1] rc1_xreal_0
a: [3,541,0,2,1,362]
s: 541
i: [1,363,0,2,1,0,0,2] rc1_xxreal_0
a: [3,546,0,2,1,363]
s: 546
d: [2,716,0,0,0,2,21,546,513]
e: 716
r: 513
s: 546
i: [1,364,0,2,1,0,0,2] rc1_xboole_0
a: [3,551,0,2,1,364]
s: 551
i: [1,368,0,2,1,0,0,255] rc2_ordinal1
a: [3,556,0,2,1,368]
s: 556
i: [1,367,0,2,1,0,0,255] rc2_ordinal1
a: [3,561,0,2,1,367]
s: 561
i: [1,366,0,2,1,0,0,255] rc2_ordinal1
a: [3,566,0,2,1,366]
s: 566
i: [1,365,0,2,1,0,0,255] rc2_ordinal1
a: [3,571,0,2,1,365]
s: 571
i: [1,370,0,2,1,0,0,255] rc5_ordinal1
a: [3,576,0,2,1,370]
s: 576
i: [1,369,0,2,1,0,0,255] rc5_ordinal1
a: [3,581,0,2,1,369]
s: 581
i: [1,372,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,586,0,2,1,372]
s: 586
i: [1,371,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,591,0,2,1,371]
s: 591
i: [1,376,0,2,1,0,0,2] rc3_xreal_0
a: [3,596,0,2,1,376]
s: 596
i: [1,375,0,2,1,0,0,2] rc3_xreal_0
a: [3,601,0,2,1,375]
s: 601
i: [1,374,0,2,1,0,0,2] rc3_xreal_0
a: [3,606,0,2,1,374]
s: 606
i: [1,373,0,2,1,0,0,2] rc3_xreal_0
a: [3,611,0,2,1,373]
s: 611
i: [1,380,0,2,1,0,0,2] rc2_xreal_0
a: [3,616,0,2,1,380]
s: 616
i: [1,379,0,2,1,0,0,2] rc2_xreal_0
a: [3,621,0,2,1,379]
s: 621
i: [1,378,0,2,1,0,0,2] rc2_xreal_0
a: [3,626,0,2,1,378]
s: 626
i: [1,377,0,2,1,0,0,2] rc2_xreal_0
a: [3,631,0,2,1,377]
s: 631
i: [1,382,0,2,1,0,0,2] rc3_xxreal_0
a: [3,636,0,2,1,382]
s: 636
i: [1,381,0,2,1,0,0,2] rc3_xxreal_0
a: [3,641,0,2,1,381]
s: 641
i: [1,384,0,2,1,0,0,2] rc2_xxreal_0
a: [3,646,0,2,1,384]
s: 646
i: [1,383,0,2,1,0,0,2] rc2_xxreal_0
a: [3,651,0,2,1,383]
s: 651
i: [1,388,0,2,1,0,0,2] rc4_xreal_0
a: [3,656,0,2,1,388]
s: 656
i: [1,387,0,2,1,0,0,2] rc4_xreal_0
a: [3,661,0,2,1,387]
s: 661
i: [1,386,0,2,1,0,0,2] rc4_xreal_0
a: [3,666,0,2,1,386]
s: 666
i: [1,385,0,2,1,0,0,2] rc4_xreal_0
a: [3,671,0,2,1,385]
s: 671
i: [1,390,0,2,1,0,0,2] rc4_xxreal_0
a: [3,676,0,2,1,390]
s: 676
i: [1,389,0,2,1,0,0,2] rc4_xxreal_0
a: [3,681,0,2,1,389]
s: 681
i: [1,395,0,2,1,0,0,2] rc1_real_1
a: [3,686,0,2,1,395]
s: 686
i: [1,394,0,2,1,0,0,2] rc1_real_1
a: [3,691,0,2,1,394]
s: 691
i: [1,393,0,2,1,0,0,2] rc1_real_1
a: [3,696,0,2,1,393]
s: 696
i: [1,392,0,2,1,0,0,2] rc1_real_1
a: [3,701,0,2,1,392]
s: 701
i: [1,340,0,3,1,0,0,255] reflexivity_r1_tarski
s: 340
i: [1,269,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,449,0,3,1,269]
s: 449
i: [1,270,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,454,0,3,1,270]
s: 454
r: 439
d: [2,718,0,2,1,2,23,439,454]
a: [3,722,0,2,1,718]
s: 722
i: [1,271,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r0_r0
a: [3,459,0,3,1,271]
s: 459
i: [1,272,0,3,1,0,0,255] t1_numerals
a: [3,464,0,3,1,272]
s: 464
i: [1,281,0,3,1,0,0,255] spc0_numerals
a: [3,499,0,3,1,281]
s: 499
i: [1,280,0,3,1,0,0,255] spc0_numerals
a: [3,504,0,3,1,280]
s: 504
i: [1,391,0,3,1,0,0,2] rc1_real_1
a: [3,706,0,3,1,391]
s: 706
d: [2,724,0,3,1,2,22,464,454]
a: [3,728,0,3,1,724]
s: 728
d: [2,730,0,3,1,2,22,504,449]
a: [3,734,0,3,1,730]
s: 734
i: [1,283,0,4,2,0,0,255] cc6_ordinal1
s: 283
i: [1,284,0,4,2,0,0,255] cc2_xcmplx_0
s: 284
d: [2,736,1,2,1,1,40,283,526]
a: [3,741,1,2,1,736]
s: 741
i: [1,285,0,4,2,0,0,255] cc2_xreal_0
s: 285
i: [1,286,0,4,2,0,0,255] cc2_xxreal_0
s: 286
d: [2,737,1,2,1,1,40,283,576]
a: [3,746,1,2,1,737]
s: 746
i: [1,287,0,4,2,0,0,255] cc1_ordinal1
s: 287
d: [2,748,1,2,1,1,40,284,526]
a: [3,753,1,2,1,748]
s: 753
i: [1,288,0,4,2,0,0,255] cc1_ordinal1
s: 288
i: [1,291,0,4,2,0,0,2] cc3_xreal_0
s: 291
d: [2,834,1,2,1,1,40,291,429]
a: [3,844,1,2,1,834]
s: 844
i: [1,292,0,4,2,0,0,2] cc4_xreal_0
s: 292
d: [2,835,1,2,1,1,40,291,424]
a: [3,849,1,2,1,835]
s: 849
i: [1,304,0,4,2,0,0,255] cc9_ordinal1
s: 304
i: [1,305,0,4,2,0,0,255] cc7_ordinal1
s: 305
d: [2,856,1,2,1,1,40,292,429]
a: [3,866,1,2,1,856]
s: 866
i: [1,306,0,4,2,0,0,255] cc3_ordinal1
s: 306
d: [2,857,1,2,1,1,40,292,424]
a: [3,871,1,2,1,857]
s: 871
i: [1,339,0,4,1,0,0,3] existence_m1_subset_1
s: 339
i: [1,273,0,4,1,0,0,255] t17_square_1
a: [3,469,0,4,1,273]
s: 469
d: [2,749,1,2,1,1,40,284,576]
a: [3,758,1,2,1,749]
s: 758
i: [1,274,0,4,1,0,0,255] dt_k5_numbers
a: [3,474,0,4,1,274]
s: 474
d: [2,760,1,2,1,1,40,285,526]
a: [3,765,1,2,1,760]
s: 765
i: [1,275,0,4,1,0,0,255] dt_k6_numbers
a: [3,479,0,4,1,275]
s: 479
i: [1,279,0,4,1,0,0,255] spc0_numerals
a: [3,509,0,4,1,279]
s: 509
d: [2,761,1,2,1,1,40,285,576]
a: [3,770,1,2,1,761]
s: 770
d: [2,968,0,4,1,2,22,474,449]
a: [3,972,0,4,1,968]
s: 972
d: [2,772,1,2,1,1,40,286,526]
a: [3,777,1,2,1,772]
s: 777
d: [2,977,0,4,1,2,22,479,449]
a: [3,981,0,4,1,977]
s: 981
d: [2,983,0,4,1,2,22,509,449]
a: [3,987,0,4,1,983]
s: 987
d: [2,773,1,2,1,1,40,286,576]
a: [3,782,1,2,1,773]
s: 782
i: [1,293,0,5,2,0,0,1] dt_k6_square_1
s: 293
d: [2,994,1,6,2,1,44,293,398]
a: [3,998,1,4,1,994]
s: 998
i: [1,294,0,5,2,0,0,1] dt_k6_square_1
s: 294
i: [1,307,0,5,2,0,0,255] t6_boole
s: 307
d: [2,784,1,2,1,1,40,287,489]
a: [3,792,1,2,1,784]
s: 792
i: [1,308,0,5,2,0,0,255] cc8_ordinal1
s: 308
d: [2,786,1,2,1,1,40,287,531]
a: [3,802,1,2,1,786]
s: 802
i: [1,309,0,5,2,0,0,2] cc1_xcmplx_0
s: 309
i: [1,310,0,5,2,0,0,3] cc1_real_1
s: 310
d: [2,809,1,2,1,1,40,288,489]
a: [3,817,1,2,1,809]
s: 817
i: [1,311,0,5,2,0,0,3] cc1_xreal_0
s: 311
d: [2,811,1,2,1,1,40,288,531]
a: [3,827,1,2,1,811]
s: 827
i: [1,336,0,5,2,0,0,255] t7_boole
s: 336
i: [1,276,0,5,1,0,0,255] rqRealAdd__k2_xcmplx_0__r0_r0_r0
a: [3,484,0,5,1,276]
s: 484
d: [2,836,1,2,1,1,40,291,541]
a: [3,854,1,2,1,836]
s: 854
a: [3,516,0,5,2,351]
s: 516
d: [2,858,1,2,1,1,40,292,541]
a: [3,876,1,2,1,858]
s: 876
d: [2,1000,0,5,2,1,22,307,454]
s: 1000
i: [1,320,0,6,3,0,0,255] cc2_ordinal1
s: 320
d: [2,1038,1,5,2,2,44,414,1000]
s: 1038
i: [1,323,0,6,3,0,0,2] cc5_xxreal_0
s: 323
d: [2,1039,1,5,2,2,44,419,1000]
s: 1039
i: [1,325,0,6,3,0,0,2] cc5_xxreal_0
s: 325
i: [1,326,0,6,3,0,0,2] cc3_xxreal_0
s: 326
d: [2,878,1,2,1,1,40,304,722]
a: [3,886,1,2,1,878]
s: 886
i: [1,328,0,6,3,0,0,2] cc3_xxreal_0
s: 328
d: [2,879,1,2,1,1,40,304,444]
a: [3,891,1,2,1,879]
s: 891
i: [1,330,0,6,3,0,0,2] cc7_xxreal_0
s: 330
i: [1,331,0,6,3,0,0,2] cc7_xxreal_0
s: 331
d: [2,880,1,2,1,1,40,304,551]
a: [3,896,1,2,1,880]
s: 896
i: [1,332,0,6,2,0,0,255] t1_subset
s: 332
d: [2,881,1,2,1,1,40,304,671]
a: [3,901,1,2,1,881]
s: 901
i: [1,341,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 341
i: [1,289,0,7,3,0,0,255] cc5_ordinal1
s: 289
d: [2,882,1,2,1,1,40,304,681]
a: [3,906,1,2,1,882]
s: 906
i: [1,290,0,7,2,0,0,1] t1_arithm
s: 290
d: [2,908,1,2,1,1,40,305,722]
a: [3,916,1,2,1,908]
s: 916
d: [2,708,0,2,1,2,23,429,402]
a: [3,712,0,2,1,708]
s: 712
i: [1,312,0,7,2,0,0,255] dt_k7_square_1
s: 312
d: [2,1137,1,4,2,2,44,712,1000]
s: 1137
i: [1,313,0,7,2,0,0,255] dt_k7_square_1
s: 313
d: [2,1150,2,2,1,3,40,1137,551]
a: [3,1157,2,2,1,1150]
s: 1157
i: [1,334,0,7,2,0,0,255] t3_subset
s: 334
i: [1,335,0,7,3,0,0,1] t8_boole
s: 335
d: [2,1152,2,2,1,3,40,1137,681]
a: [3,1162,2,2,1,1152]
s: 1162
i: [1,282,0,8,3,0,0,255] cc10_ordinal1
s: 282
d: [2,1164,3,2,1,1,40,1157,292]
a: [3,1169,3,2,1,1164]
s: 1169
i: [1,314,0,8,2,0,0,255] redefinition_k7_square_1
s: 314
i: [1,315,0,8,2,0,0,255] redefinition_k7_square_1
s: 315
d: [2,1165,3,2,1,1,40,1157,291]
a: [3,1174,3,2,1,1165]
s: 1174
i: [1,317,0,8,4,0,0,2] cc4_xxreal_0
s: 317
d: [2,1210,2,6,3,1,40,317,866]
a: [3,1227,2,2,1,1210]
s: 1227
d: [2,1299,3,0,0,3,52,722,403,1227,335]
e: 1299
d: [2,1300,3,0,0,3,52,722,403,1227,335]
e: 1300
d: [2,1301,3,0,0,3,52,403,1227,1000]
e: 1301
r: 403
r: 1227
i: [1,319,0,8,4,0,0,2] cc6_xxreal_0
s: 319
i: [1,321,0,8,4,0,0,2] cc8_xxreal_0
s: 321
a: [3,1226,2,2,1,1210]
s: 1226
i: [1,333,0,8,3,0,0,255] t2_subset
s: 333
a: [3,1231,2,2,1,1210]
s: 1231
i: [1,348,0,8,3,0,0,1] fc2_xcmplx_0
s: 348
i: [1,350,0,8,3,0,0,1] fc5_xreal_0
s: 350
d: [2,1371,2,0,0,3,52,429,424,998,350]
e: 1371
d: [2,1372,2,4,2,1,40,350,998]
d: [2,1379,2,2,1,2,21,1372,424]
d: [2,1380,2,0,0,3,21,1379,429]
e: 1380
r: 998
d: [2,1211,2,6,3,1,40,317,871]
a: [3,1240,2,2,1,1211]
s: 1240
d: [2,1382,3,0,0,3,52,722,407,1240,335]
e: 1382
d: [2,1383,3,0,0,3,52,722,407,1240,335]
e: 1383
d: [2,1384,3,0,0,3,52,407,1240,1000]
e: 1384
r: 407
r: 1240
s: 403
a: [3,1239,2,2,1,1211]
s: 1239
i: [1,338,0,9,3,0,0,255] t5_subset
s: 338
i: [1,337,0,10,3,0,0,255] t4_subset
s: 337
a: [3,1244,2,2,1,1211]
s: 1244
i: [1,352,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 352
d: [2,1178,3,2,1,1,40,1162,291]
a: [3,1182,3,2,1,1178]
s: 1182
i: [1,298,0,11,5,0,0,1] t2_real
s: 298
i: [1,299,0,11,5,0,0,1] t1_real
s: 299
a: [3,1234,2,2,1,1210]
s: 1234
i: [1,300,0,11,5,0,0,1] t8_real
s: 300
a: [3,1247,2,2,1,1211]
s: 1247
i: [1,301,0,11,5,0,0,1] t7_real
s: 301
i: [1,302,0,11,5,0,0,1] t3_real
s: 302
a: [3,997,1,4,1,994]
s: 997
i: [1,303,0,11,5,0,0,1] t4_real
s: 303
d: [2,1461,2,4,1,1,40,997,292]
a: [3,1466,2,4,1,1461]
s: 1466
i: [1,349,0,11,3,0,0,1] commutativity_k2_xcmplx_0
s: 349
i: [1,295,0,12,4,0,0,1] t24_square_1
s: 295
d: [2,1462,2,4,1,1,40,997,291]
a: [3,1471,2,4,1,1462]
s: 1471
i: [1,342,0,12,5,0,0,1] fc9_xreal_0
s: 342
d: [2,1473,3,12,3,1,40,1466,317]
a: [3,1477,3,4,1,1473]
s: 1477
r: 1477
i: [1,347,0,12,5,0,0,1] fc10_xreal_0
s: 347
i: [1,353,0,12,5,0,0,1] fc13_xreal_0
s: 353
a: [3,1476,3,4,1,1473]
s: 1476
i: [1,354,0,12,5,0,0,1] fc14_xreal_0
s: 354
d: [2,1416,1,8,4,1,40,298,414]
d: [2,1425,1,6,3,2,21,1416,424]
d: [2,1426,1,4,2,3,21,1425,712]
a: [3,1429,1,2,1,1426]
s: 1429
i: [1,355,0,12,5,0,0,1] fc12_xreal_0
s: 355
i: [1,356,0,12,5,0,0,1] fc11_xreal_0
s: 356
d: [2,1581,1,12,6,0,44,355,290]
d: [2,1586,1,10,5,1,21,1581,712]
d: [2,1587,1,8,4,1,21,1586,291]
a: [3,1591,1,2,1,1587]
s: 1591
i: [1,296,0,13,6,0,0,1] t5_real
s: 296
d: [2,1577,2,4,2,2,44,1429,1000]
s: 1577
r: 323
i: [1,297,0,13,6,0,0,1] t6_real
s: 297
i: [1,346,0,14,4,0,0,255] existence_m2_subset_1
s: 346
d: [2,1605,2,4,2,2,44,1591,1000]
s: 1605
r: 326
i: [1,343,0,15,5,0,0,255] dt_m2_subset_1
s: 343
a: [3,1481,3,4,1,1473]
s: 1481
i: [1,344,0,15,5,0,0,255] redefinition_m2_subset_1
s: 344
i: [1,345,0,15,5,0,0,255] redefinition_m2_subset_1
s: 345
d: [2,1508,1,17,4,1,44,295,398]
d: [2,1509,1,14,3,1,19,1508]
d: [2,1510,1,10,2,2,21,1509,997]
a: [3,1514,1,5,1,1510]
s: 1514
r: 1476
r: 1481
r: 1514
i: [1,357,0,17,4,0,0,1] spc6_arithm
s: 357
s: 1477
a: [3,1513,1,5,1,1510]
s: 1513
a: [3,1518,1,5,1,1510]
s: 1518
r: 398
r: 997
r: 1466
r: 1471
r: 1477
r: 1513
d: [2,1768,2,10,5,1,44,355,1518]
d: [2,1782,2,8,4,2,21,1768,1591]
d: [2,1783,2,6,3,3,21,1782,429]
d: [2,1784,2,4,2,4,21,1783,1234]
d: [2,1785,2,2,1,5,21,1784,424]
d: [2,1786,2,0,0,6,21,1785,1244]
e: 1786
d: [2,1769,2,10,5,1,44,356,1518]
d: [2,1788,2,8,4,2,21,1769,1591]
d: [2,1789,2,6,3,3,21,1788,424]
d: [2,1790,2,4,2,4,21,1789,1247]
d: [2,1791,2,2,1,5,21,1790,429]
d: [2,1792,2,0,0,6,21,1791,1231]
e: 1792
r: 1244
r: 1247
r: 1577
r: 1429
a: [3,1430,1,2,1,1426]
s: 1430
s: 323
r: 1430
r: 1605
r: 1591
s: 1429
s: 326
a: [3,1590,1,2,1,1587]
s: 1590
r: 403
r: 1226
r: 1231
r: 1239
r: 1518
r: 1590
s: 1591
s: 402
r: 419
r: 429
r: 844
r: 866
r: 1039
r: 1234
s: 1227
s: 407
s: 398
a: [3,1230,2,2,1,1210]
s: 1230
d: [2,1900,0,6,1,2,22,398,402]
a: [3,1904,0,6,1,1900]
s: 1904
s: 1244
d: [2,1761,1,3,1,2,23,1513,1518]
a: [3,1777,1,3,1,1761]
s: 1777
s: 997
s: 1239
d: [2,1908,1,17,4,1,44,295,1904]
d: [2,1910,1,14,3,1,19,1908]
a: [3,1913,1,4,1,1910]
s: 1913
s: 1247
s: 1514
a: [3,1517,1,5,1,1510]
s: 1517
d: [2,1305,3,2,1,1,40,1227,304]
a: [3,1312,3,2,1,1305]
s: 1312
d: [2,1894,1,5,2,3,23,1039,402]
s: 1894
d: [2,1304,3,2,1,1,40,1227,305]
a: [3,1317,3,2,1,1304]
s: 1317
a: [3,1918,1,5,1,1910]
s: 1918
a: [3,1921,1,5,1,1910]
s: 1921
d: [2,1303,3,2,1,1,40,1227,306]
a: [3,1322,3,2,1,1303]
s: 1322
d: [2,1566,1,12,6,0,44,354,290]
d: [2,1571,1,10,5,1,21,1566,712]
d: [2,1572,1,8,4,1,21,1571,291]
a: [3,1575,1,6,3,1572]
s: 1575
d: [2,1986,2,0,0,4,52,1429,1591,712,1575]
e: 1986
d: [2,1987,2,4,2,2,40,1575,712]
d: [2,2002,2,2,1,3,21,1987,1591]
d: [2,2003,2,0,0,4,21,2002,1429]
e: 2003
r: 1575
d: [2,1388,3,2,1,1,40,1240,304]
a: [3,1395,3,2,1,1388]
s: 1395
d: [2,1907,1,8,2,2,44,1904,1000]
s: 1907
d: [2,909,1,2,1,1,40,305,444]
a: [3,921,1,2,1,909]
s: 921
d: [2,1387,3,2,1,1,40,1240,305]
a: [3,1400,3,2,1,1387]
s: 1400
d: [2,910,1,2,1,1,40,305,551]
a: [3,926,1,2,1,910]
s: 926
d: [2,1386,3,2,1,1,40,1240,306]
a: [3,1405,3,2,1,1386]
s: 1405
r: 1400
r: 1405
d: [2,911,1,2,1,1,40,305,671]
a: [3,931,1,2,1,911]
s: 931
d: [2,912,1,2,1,1,40,305,681]
a: [3,936,1,2,1,912]
s: 936
a: [3,1399,3,2,1,1387]
s: 1399
d: [2,938,1,2,1,1,40,306,722]
a: [3,946,1,2,1,938]
s: 946
a: [3,1404,3,2,1,1386]
s: 1404
d: [2,939,1,2,1,1,40,306,444]
a: [3,951,1,2,1,939]
s: 951
d: [2,940,1,2,1,1,40,306,551]
a: [3,956,1,2,1,940]
s: 956
d: [2,1421,2,10,5,2,40,298,1039]
d: [2,1437,2,8,4,3,21,1421,429]
d: [2,1438,2,6,3,4,21,1437,1137]
a: [3,1441,2,4,2,1438]
s: 1441
d: [2,2095,3,0,0,3,52,1429,722,1441]
e: 2095
d: [2,2096,3,2,1,2,40,1441,722]
d: [2,2102,3,0,0,3,21,2096,1429]
e: 2102
r: 1441
d: [2,941,1,2,1,1,40,306,671]
a: [3,961,1,2,1,941]
s: 961
s: 1466
d: [2,942,1,2,1,1,40,306,681]
a: [3,966,1,2,1,942]
s: 966
d: [2,1011,1,2,1,1,40,309,499]
a: [3,1017,1,2,1,1011]
s: 1017
s: 1471
d: [2,1024,1,2,1,1,40,310,499]
a: [3,1030,1,2,1,1024]
s: 1030
d: [2,1452,2,10,5,2,40,299,1038]
d: [2,1834,2,8,4,3,21,1452,424]
d: [2,1835,2,6,3,4,21,1834,1137]
a: [3,1838,2,4,2,1835]
s: 1838
r: 326
d: [2,1047,1,4,2,1,40,323,601]
d: [2,1049,1,2,1,2,21,1047,606]
a: [3,1053,1,2,1,1049]
s: 1053
d: [2,1048,1,4,2,1,40,323,636]
d: [2,1055,1,2,1,2,21,1048,641]
a: [3,1059,1,2,1,1055]
s: 1059
d: [2,1850,2,4,2,2,44,1429,1000]
s: 1850
r: 323
d: [2,1061,1,4,2,1,40,325,601]
d: [2,1063,1,2,1,2,21,1061,606]
a: [3,1067,1,2,1,1063]
s: 1067
d: [2,1939,2,4,1,1,40,1913,292]
a: [3,1945,2,4,1,1939]
s: 1945
d: [2,1062,1,4,2,1,40,325,636]
d: [2,1069,1,2,1,2,21,1062,641]
a: [3,1073,1,2,1,1069]
s: 1073
d: [2,1075,1,4,2,1,40,326,621]
d: [2,1078,1,2,1,2,21,1075,626]
a: [3,1082,1,2,1,1078]
s: 1082
d: [2,1940,2,4,1,1,40,1913,291]
a: [3,1950,2,4,1,1940]
s: 1950
d: [2,1076,1,4,2,1,40,326,646]
d: [2,1084,1,2,1,2,21,1076,651]
a: [3,1088,1,2,1,1084]
s: 1088
d: [2,1967,2,12,4,1,40,1918,300]
d: [2,1970,2,10,3,2,21,1967,1591]
d: [2,1971,2,6,2,3,21,1970,1913]
d: [2,1972,2,4,1,4,21,1971,712]
a: [3,1976,2,4,1,1972]
s: 1976
d: [2,2160,3,0,0,5,52,1429,712,1247,424,1976,342]
e: 2160
d: [2,2161,3,8,4,1,40,1976,342]
d: [2,2166,3,6,3,2,21,2161,424]
d: [2,2167,3,4,2,3,21,2166,1247]
d: [2,2168,3,2,1,4,21,2167,712]
d: [2,2169,3,0,0,5,21,2168,1429]
e: 2169
r: 407
r: 1239
r: 1244
r: 1399
r: 1404
r: 1907
r: 1904
r: 1918
r: 1976
s: 406
r: 414
r: 424
r: 849
r: 871
r: 1038
r: 1247
r: 1395
r: 1913
r: 1921
r: 1945
r: 1950
a: [3,1903,0,6,1,1900]
s: 1903
s: 1240
d: [2,2212,0,6,1,2,22,1903,406]
a: [3,2216,0,6,1,2212]
s: 2216
r: 1894
r: 402
r: 406
r: 2216
s: 424
s: 429
s: 403
s: 844
s: 407
s: 849
s: 414
s: 419
s: 866
s: 398
s: 871
s: 1903
a: [3,2215,0,6,1,2212]
s: 2215
s: 1226
s: 997
s: 1231
s: 1913
d: [2,2186,1,4,1,2,23,1913,406]
a: [3,2195,1,4,1,2186]
s: 2195
s: 1234
s: 1038
s: 1239
s: 1039
s: 1514
s: 1244
s: 1517
s: 1247
a: [3,1917,1,5,1,1910]
s: 1917
s: 1921
s: 1312
d: [2,2187,1,5,1,2,23,1921,406]
a: [3,2200,1,5,1,2187]
s: 2200
r: 2200
s: 1317
a: [3,2199,1,5,1,2187]
s: 2199
r: 2195
r: 2215
d: [2,2234,1,5,2,2,44,414,1000]
s: 2234
s: 1322
r: 1317
r: 1322
d: [2,2247,1,5,2,2,44,419,1000]
s: 2247
a: [3,1316,3,2,1,1304]
s: 1316
d: [2,2264,1,8,2,2,44,1903,1000]
s: 2264
d: [2,2265,1,8,2,2,44,2215,1000]
s: 2265
a: [3,1321,3,2,1,1303]
s: 1321
d: [2,1077,1,4,2,1,40,326,691]
d: [2,1090,1,2,1,2,21,1077,696]
a: [3,1094,1,2,1,1090]
s: 1094
s: 1395
d: [2,1101,1,4,2,1,40,289,499]
a: [3,1110,1,2,1,1101]
s: 1110
d: [2,1184,1,4,2,1,40,282,972]
a: [3,1190,1,2,1,1184]
s: 1190
s: 1400
d: [2,1214,1,6,3,1,40,317,546]
a: [3,1279,1,2,1,1214]
s: 1279
d: [2,2415,2,2,1,3,40,1279,1137]
a: [3,2422,2,2,1,2415]
s: 2422
d: [2,1003,1,3,1,0,40,308,339]
a: [3,1009,1,3,1,1003]
s: 1009
d: [2,1013,1,3,1,0,40,309,339]
a: [3,1022,1,3,1,1013]
s: 1022
s: 1405
d: [2,1026,1,3,1,0,40,310,339]
a: [3,1035,1,3,1,1026]
s: 1035
d: [2,2440,3,2,1,1,40,2422,291]
a: [3,2444,3,2,1,2440]
s: 2444
d: [2,1196,1,7,2,1,44,314,469]
a: [3,1202,1,3,1,1196]
s: 1202
d: [2,1327,1,5,2,1,40,333,728]
d: [2,1341,1,3,1,2,21,1327,494]
a: [3,1345,1,3,1,1341]
s: 1345
d: [2,2036,4,2,1,1,40,1405,288]
a: [3,2041,4,2,1,2036]
s: 2041
d: [2,1328,1,5,2,1,40,333,734]
d: [2,1347,1,3,1,2,21,1328,494]
a: [3,1351,1,3,1,1347]
s: 1351
d: [2,2037,4,2,1,1,40,1405,287]
a: [3,2046,4,2,1,2037]
s: 2046
d: [2,1329,1,5,2,1,40,333,499]
d: [2,1353,1,3,1,2,21,1329,434]
a: [3,1357,1,3,1,1353]
s: 1357
d: [2,1330,1,5,2,1,40,333,706]
d: [2,1359,1,3,1,2,21,1330,434]
a: [3,1363,1,3,1,1359]
s: 1363
s: 1466
a: [3,1207,1,4,1,1196]
s: 1207
s: 1471
d: [2,1325,1,7,2,1,40,333,972]
a: [3,1335,1,4,1,1325]
s: 1335
a: [3,1338,1,3,1,1325]
s: 1338
s: 1945
d: [2,992,1,5,2,0,40,293,292]
s: 992
s: 1950
d: [2,993,1,5,2,0,40,293,291]
s: 993
d: [2,1042,1,5,2,2,44,728,1000]
s: 1042
a: [3,1975,2,4,1,1972]
s: 1975
d: [2,1097,1,5,2,0,40,289,339]
s: 1097
d: [2,2188,2,4,1,2,23,1945,406]
a: [3,2205,2,4,1,2188]
s: 2205
d: [2,1040,1,6,2,2,44,469,1000]
s: 1040
d: [2,1043,1,6,2,2,44,981,1000]
s: 1043
d: [2,2189,2,4,1,2,23,1950,406]
a: [3,2210,2,4,1,2189]
s: 2210
d: [2,1140,1,6,2,0,40,312,310]
s: 1140
d: [2,2392,2,6,2,3,44,2265,484]
d: [2,2403,2,4,1,4,21,2392,444]
a: [3,2407,2,4,1,2403]
s: 2407
d: [2,1141,1,6,2,0,40,312,309]
s: 1141
d: [2,1186,1,6,2,0,40,282,339]
s: 1186
s: 1477
d: [2,2624,4,0,0,3,52,722,1517,1477,335]
e: 2624
d: [2,2625,4,0,0,3,52,722,1517,1477,335]
e: 2625
d: [2,2626,4,0,0,3,52,1517,1477,1000]
e: 2626
r: 1477
d: [2,1324,1,6,2,0,40,333,339]
s: 1324
s: 1476
d: [2,1410,1,6,2,1,40,337,972]
s: 1410
d: [2,1037,1,7,3,1,44,1000,1000]
s: 1037
a: [3,1485,3,4,1,1473]
s: 1485
d: [2,2646,4,0,0,5,52,1234,429,1247,424,1485,342]
e: 2646
d: [2,2647,4,8,4,1,40,1485,342]
d: [2,2651,4,6,3,2,21,2647,424]
d: [2,2652,4,4,2,3,21,2651,1247]
d: [2,2653,4,2,1,4,21,2652,429]
d: [2,2654,4,0,0,5,21,2653,1234]
e: 2654
i: [1,264,0,9,2,1,0,0] t31_square_1
i: [1,263,0,9,2,1,0,0] t31_square_1
d: [2,1001,1,2,1,1,40,308,728]
d: [2,1118,2,2,1,1,40,916,286]
d: [2,1119,2,2,1,1,40,916,285]
d: [2,1120,2,2,1,1,40,916,284]
d: [2,1197,1,7,2,1,44,469,314]
d: [2,1649,1,11,4,1,40,343,981]
d: [2,1653,1,7,3,2,21,1649,972]
d: [2,1654,1,5,2,3,21,1653,494]
d: [2,1655,1,3,1,4,21,1654,434]
d: [2,1820,2,4,2,1,40,323,1430]
d: [2,1825,2,2,1,2,21,1820,722]
d: [2,1840,2,8,4,2,21,1769,424]
d: [2,1841,2,6,3,3,21,1840,429]
d: [2,1842,2,4,2,4,21,1841,1231]
d: [2,1844,2,8,4,2,21,1768,429]
d: [2,1845,2,6,3,3,21,1844,1234]
d: [2,1846,2,4,2,4,21,1845,424]
d: [2,1863,2,4,2,1,40,1590,326]
d: [2,1865,2,2,1,2,21,1863,722]
a: [3,399,0,6,1,265]
d: [2,2177,0,6,1,2,22,399,402]
d: [2,2220,1,6,2,1,44,2216,290]
d: [2,2222,1,2,1,1,40,424,292]
d: [2,2226,1,2,1,1,40,429,292]
d: [2,2233,1,8,4,1,40,414,298]
d: [2,2239,1,6,3,2,21,2233,1429]
d: [2,2240,1,4,2,3,21,2239,424]
d: [2,2241,1,2,1,4,21,2240,712]
d: [2,2246,1,8,4,1,40,419,298]
d: [2,2252,1,6,3,2,21,2246,1429]
d: [2,2253,1,4,2,3,21,2252,429]
d: [2,2254,1,2,1,4,21,2253,712]
d: [2,2256,2,6,3,1,40,866,317]
d: [2,2258,1,17,4,1,44,295,398]
d: [2,2260,1,14,3,1,19,2258]
d: [2,2262,2,6,3,1,40,871,317]
d: [2,2311,2,12,4,1,40,1514,300]
d: [2,2314,2,10,3,2,21,2311,1591]
d: [2,2315,2,6,2,3,21,2314,997]
d: [2,2316,2,4,1,4,21,2315,712]
f: 2654,2316,2262,2260,2256,2254,2241,2226,2222,2220,2212,2177,1865,1846,1842,1825,1655,1384,1380,1301,1197,1120,1119,1118,1001,977,968,724,718,277,275,274,273,272,270,269,267,266,259,260,261,262,263,264,265
% Refutation found. Thanks to Tanya!
% SZS status Theorem for square_1__t31_square_1
% SZS output start Proof for square_1__t31_square_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0)) => (k6_numbers = k6_square_1(k2_xcmplx_0(X0,X1)) <=> (k6_numbers = X1 & k6_numbers = X0))))) [input t31_square_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0)) => (k6_numbers = k6_square_1(k2_xcmplx_0(X0,X1)) <=> (k6_numbers = X1 & k6_numbers = X0))))) [negated conjecture 1]
11. ! [X0] : (v7_ordinal1(X0) => v1_xcmplx_0(X0)) [input cc2_xcmplx_0]
12. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
13. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
14. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
15. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) => (v2_xxreal_0(X0) & v1_xxreal_0(X0))) [input cc4_xxreal_0]
17. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
20. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
23. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & v2_xxreal_0(X0)) => v2_xxreal_0(k2_xcmplx_0(X0,X1))) [input fc11_xreal_0]
24. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & v2_xxreal_0(X0)) => v2_xxreal_0(k2_xcmplx_0(X1,X0))) [input fc12_xreal_0]
27. v3_ordinal1(k4_ordinal1) & ~v1_xboole_0(k4_ordinal1) [input fc6_ordinal1]
28. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k2_xcmplx_0(X0,X1))) [input fc9_xreal_0]
37. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
40. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
49. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
52. k4_ordinal1 = k5_numbers [input redefinition_k5_numbers]
56. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [input dt_k5_numbers]
58. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ! [X2] : (m2_subset_1(X2,X0,X1) => m1_subset_1(X2,X0))) [input dt_m2_subset_1]
67. ~v1_xboole_0(k1_numbers) [input fc1_numbers]
68. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
79. ! [X0] : (v1_xcmplx_0(X0) => k2_xcmplx_0(X0,k6_numbers) = X0) [input t1_arithm]
81. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
87. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
89. ! [X0] : (m1_subset_1(X0,k1_numbers) => k7_square_1(X0) = k6_square_1(X0)) [input redefinition_k7_square_1]
91. m2_subset_1(k6_numbers,k1_numbers,k5_numbers) [input dt_k6_numbers]
97. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
98. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k2_xcmplx_0(X0,X1))) [input fc5_xreal_0]
102. k6_numbers = k7_square_1(k6_numbers) [input t17_square_1]
103. ! [X0] : (v1_xreal_0(X0) => ((k6_numbers = k6_square_1(X0) & r1_xxreal_0(k6_numbers,X0)) => k6_numbers = X0)) [input t24_square_1]
117. ? [X0] : (? [X1] : (((k6_numbers = k6_square_1(k2_xcmplx_0(X0,X1)) <~> (k6_numbers = X1 & k6_numbers = X0)) & (r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
118. ? [X0] : (? [X1] : ((k6_numbers = k6_square_1(k2_xcmplx_0(X0,X1)) <~> (k6_numbers = X1 & k6_numbers = X0)) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 117]
121. ! [X0] : (v1_xcmplx_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
122. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
123. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 13]
126. ! [X0] : (k2_xcmplx_0(X0,k6_numbers) = X0 | ~v1_xcmplx_0(X0)) [ennf transformation 79]
128. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 97]
131. ! [X0] : ((k6_numbers = X0 | (k6_numbers != k6_square_1(X0) | ~r1_xxreal_0(k6_numbers,X0))) | ~v1_xreal_0(X0)) [ennf transformation 103]
132. ! [X0] : (k6_numbers = X0 | k6_numbers != k6_square_1(X0) | ~r1_xxreal_0(k6_numbers,X0) | ~v1_xreal_0(X0)) [flattening 131]
137. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 40]
138. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 137]
141. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 49]
142. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 141]
152. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 81]
153. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 20]
159. ! [X0] : (k7_square_1(X0) = k6_square_1(X0) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 89]
161. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | (v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0))) [ennf transformation 15]
162. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0)) [flattening 161]
169. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 17]
170. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 169]
171. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 14]
172. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 171]
185. ! [X0,X1] : (~v3_xxreal_0(k2_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 28]
186. ! [X0,X1] : (~v3_xxreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 185]
187. ! [X0,X1] : (! [X2] : (m1_subset_1(X2,X0) | ~m2_subset_1(X2,X0,X1)) | (~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 58]
188. ! [X0,X1] : (! [X2] : (m1_subset_1(X2,X0) | ~m2_subset_1(X2,X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 187]
199. ! [X0,X1] : (v1_xreal_0(k2_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 98]
200. ! [X0,X1] : (v1_xreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 199]
209. ! [X0,X1] : (v2_xxreal_0(k2_xcmplx_0(X1,X0)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0))) [ennf transformation 24]
210. ! [X0,X1] : (v2_xxreal_0(k2_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0)) [flattening 209]
211. ! [X0,X1] : (v2_xxreal_0(k2_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0))) [ennf transformation 23]
212. ! [X0,X1] : (v2_xxreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0)) [flattening 211]
215. ? [X0] : (? [X1] : ((((k6_numbers != X1 | k6_numbers != X0) | k6_numbers != k6_square_1(k2_xcmplx_0(X0,X1))) & ((k6_numbers = X1 & k6_numbers = X0) | k6_numbers = k6_square_1(k2_xcmplx_0(X0,X1)))) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [nnf transformation 118]
216. ? [X0] : (? [X1] : ((k6_numbers != X1 | k6_numbers != X0 | k6_numbers != k6_square_1(k2_xcmplx_0(X0,X1))) & ((k6_numbers = X1 & k6_numbers = X0) | k6_numbers = k6_square_1(k2_xcmplx_0(X0,X1))) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 215]
217. ? [X0] : (? [X1] : ((k6_numbers != X1 | k6_numbers != X0 | k6_numbers != k6_square_1(k2_xcmplx_0(X0,X1))) & ((k6_numbers = X1 & k6_numbers = X0) | k6_numbers = k6_square_1(k2_xcmplx_0(X0,X1))) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : ((k6_numbers != X1 | k6_numbers != sK0 | k6_numbers != k6_square_1(k2_xcmplx_0(sK0,X1))) & ((k6_numbers = X1 & k6_numbers = sK0) | k6_numbers = k6_square_1(k2_xcmplx_0(sK0,X1))) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
218. ? [X1] : ((k6_numbers != X1 | k6_numbers != sK0 | k6_numbers != k6_square_1(k2_xcmplx_0(sK0,X1))) & ((k6_numbers = X1 & k6_numbers = sK0) | k6_numbers = k6_square_1(k2_xcmplx_0(sK0,X1))) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) => ((k6_numbers != sK1 | k6_numbers != sK0 | k6_numbers != k6_square_1(k2_xcmplx_0(sK0,sK1))) & ((k6_numbers = sK1 & k6_numbers = sK0) | k6_numbers = k6_square_1(k2_xcmplx_0(sK0,sK1))) & r1_xxreal_0(k6_numbers,sK1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) [choice axiom]
219. ((k6_numbers != sK1 | k6_numbers != sK0 | k6_numbers != k6_square_1(k2_xcmplx_0(sK0,sK1))) & ((k6_numbers = sK1 & k6_numbers = sK0) | k6_numbers = k6_square_1(k2_xcmplx_0(sK0,sK1))) & r1_xxreal_0(k6_numbers,sK1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 216,218,217]
259. v1_xreal_0(sK0) [cnf transformation 219] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
260. v1_xreal_0(sK1) [cnf transformation 219] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
261. r1_xxreal_0(k6_numbers,sK0) [cnf transformation 219] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
262. r1_xxreal_0(k6_numbers,sK1) [cnf transformation 219] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
263. k6_numbers = sK0 | k6_numbers = k6_square_1(k2_xcmplx_0(sK0,sK1)) [cnf transformation 219] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
264. k6_numbers = sK1 | k6_numbers = k6_square_1(k2_xcmplx_0(sK0,sK1)) [cnf transformation 219] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
265. k6_numbers != sK1 | k6_numbers != sK0 | k6_numbers != k6_square_1(k2_xcmplx_0(sK0,sK1)) [cnf transformation 219] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
266. ~v1_xboole_0(k1_numbers) [cnf transformation 67] {a:0,w:2,wCS:800,sine:3,thAx:0,allAx:1,thDist:-1}
267. v1_xboole_0(k1_xboole_0) [cnf transformation 68] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
269. k4_ordinal1 = k5_numbers [cnf transformation 52] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
270. k6_numbers = k1_xboole_0 [cnf transformation 87] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
272. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 37] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
273. k6_numbers = k7_square_1(k6_numbers) [cnf transformation 102] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
274. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [cnf transformation 56] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
275. m2_subset_1(k6_numbers,k1_numbers,k5_numbers) [cnf transformation 91] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
277. ~v1_xboole_0(k4_ordinal1) [cnf transformation 27] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
284. ~v7_ordinal1(X0) | v1_xcmplx_0(X0) [cnf transformation 121] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
285. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 122] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
286. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 123] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
290. k2_xcmplx_0(X0,k6_numbers) = X0 | ~v1_xcmplx_0(X0) [cnf transformation 126] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
292. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 128] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
295. k6_numbers != k6_square_1(X0) | k6_numbers = X0 | ~r1_xxreal_0(k6_numbers,X0) | ~v1_xreal_0(X0) [cnf transformation 132] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
298. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 138] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
300. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 142] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
307. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 152] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
308. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 153] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
314. k7_square_1(X0) = k6_square_1(X0) | ~m1_subset_1(X0,k1_numbers) [cnf transformation 159] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
317. ~v1_xxreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xboole_0(X0) [cnf transformation 162] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
323. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 170] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
326. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 172] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
342. ~v3_xxreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 186] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
343. ~m2_subset_1(X2,X0,X1) | m1_subset_1(X2,X0) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 188] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
350. v1_xreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 200] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
355. v2_xxreal_0(k2_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0) [cnf transformation 210] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
356. v2_xxreal_0(k2_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | ~v2_xxreal_0(X0) [cnf transformation 212] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
397. 1 <=> k6_numbers = k6_square_1(k2_xcmplx_0(sK0,sK1)) [avatar definition]
398. k6_numbers = k6_square_1(k2_xcmplx_0(sK0,sK1)) <- (1) [avatar component clause 397] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
399. k6_numbers != k6_square_1(k2_xcmplx_0(sK0,sK1)) <- (~1) [avatar component clause 397] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
401. 2 <=> k6_numbers = sK0 [avatar definition]
402. k6_numbers = sK0 <- (2) [avatar component clause 401] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
403. k6_numbers != sK0 <- (~2) [avatar component clause 401] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
405. 3 <=> k6_numbers = sK1 [avatar definition]
406. k6_numbers = sK1 <- (3) [avatar component clause 405] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
407. k6_numbers != sK1 <- (~3) [avatar component clause 405] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
408. ~1 | ~2 | ~3 [avatar split clause 265,405,401,397]
409. 1 | 3 [avatar split clause 264,405,397]
410. 1 | 2 [avatar split clause 263,401,397]
412. 4 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
414. r1_xxreal_0(k6_numbers,sK1) <- (4) [avatar component clause 412] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
415. 4 [avatar split clause 262,412]
417. 5 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
419. r1_xxreal_0(k6_numbers,sK0) <- (5) [avatar component clause 417] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
420. 5 [avatar split clause 261,417]
422. 6 <=> v1_xreal_0(sK1) [avatar definition]
424. v1_xreal_0(sK1) <- (6) [avatar component clause 422] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
425. 6 [avatar split clause 260,422]
427. 7 <=> v1_xreal_0(sK0) [avatar definition]
429. v1_xreal_0(sK0) <- (7) [avatar component clause 427] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
430. 7 [avatar split clause 259,427]
432. 8 <=> v1_xboole_0(k1_numbers) [avatar definition]
434. ~v1_xboole_0(k1_numbers) <- (~8) [avatar component clause 432] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
435. ~8 [avatar split clause 266,432]
437. 9 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
439. v1_xboole_0(k1_xboole_0) <- (9) [avatar component clause 437] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
440. 9 [avatar split clause 267,437]
447. 11 <=> k4_ordinal1 = k5_numbers [avatar definition]
449. k4_ordinal1 = k5_numbers <- (11) [avatar component clause 447] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
450. 11 [avatar split clause 269,447]
452. 12 <=> k6_numbers = k1_xboole_0 [avatar definition]
454. k6_numbers = k1_xboole_0 <- (12) [avatar component clause 452] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
455. 12 [avatar split clause 270,452]
462. 14 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
464. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (14) [avatar component clause 462] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
465. 14 [avatar split clause 272,462]
467. 15 <=> k6_numbers = k7_square_1(k6_numbers) [avatar definition]
469. k6_numbers = k7_square_1(k6_numbers) <- (15) [avatar component clause 467] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
470. 15 [avatar split clause 273,467]
472. 16 <=> m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [avatar definition]
474. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) <- (16) [avatar component clause 472] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
475. 16 [avatar split clause 274,472]
477. 17 <=> m2_subset_1(k6_numbers,k1_numbers,k5_numbers) [avatar definition]
479. m2_subset_1(k6_numbers,k1_numbers,k5_numbers) <- (17) [avatar component clause 477] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
480. 17 [avatar split clause 275,477]
492. 20 <=> v1_xboole_0(k4_ordinal1) [avatar definition]
494. ~v1_xboole_0(k4_ordinal1) <- (~20) [avatar component clause 492] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
495. ~20 [avatar split clause 277,492]
710. 64 <=> v1_xreal_0(k6_numbers) [avatar definition]
712. v1_xreal_0(k6_numbers) <- (64) [avatar component clause 710] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
718. v1_xboole_0(k6_numbers) <- (9, 12) [backward demodulation 439,454] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
720. 65 <=> v1_xboole_0(k6_numbers) [avatar definition]
722. v1_xboole_0(k6_numbers) <- (65) [avatar component clause 720] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
723. 65 | ~9 | ~12 [avatar split clause 718,452,437,720]
724. m1_subset_1(k6_numbers,k4_ordinal1) <- (12, 14) [forward demodulation 464,454] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
726. 66 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
728. m1_subset_1(k6_numbers,k4_ordinal1) <- (66) [avatar component clause 726] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
729. 66 | ~12 | ~14 [avatar split clause 724,462,452,726]
864. 87 <=> v1_xxreal_0(sK0) [avatar definition]
866. v1_xxreal_0(sK0) <- (87) [avatar component clause 864] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
869. 88 <=> v1_xxreal_0(sK1) [avatar definition]
871. v1_xxreal_0(sK1) <- (88) [avatar component clause 869] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
914. 95 <=> v7_ordinal1(k6_numbers) [avatar definition]
916. v7_ordinal1(k6_numbers) <- (95) [avatar component clause 914] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
968. m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) <- (11, 16) [forward demodulation 474,449] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
970. 105 <=> m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) [avatar definition]
972. m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) <- (105) [avatar component clause 970] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
973. 105 | ~11 | ~16 [avatar split clause 968,472,447,970]
977. m2_subset_1(k6_numbers,k1_numbers,k4_ordinal1) <- (11, 17) [forward demodulation 479,449] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
979. 106 <=> m2_subset_1(k6_numbers,k1_numbers,k4_ordinal1) [avatar definition]
981. m2_subset_1(k6_numbers,k1_numbers,k4_ordinal1) <- (106) [avatar component clause 979] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
982. 106 | ~11 | ~17 [avatar split clause 977,477,447,979]
996. 108 <=> v1_xreal_0(k2_xcmplx_0(sK0,sK1)) [avatar definition]
997. v1_xreal_0(k2_xcmplx_0(sK0,sK1)) <- (108) [avatar component clause 996] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
998. ~v1_xreal_0(k2_xcmplx_0(sK0,sK1)) <- (~108) [avatar component clause 996] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1000. k6_numbers = X0 | ~v1_xboole_0(X0) <- (12) [forward demodulation 307,454] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1001. v7_ordinal1(k6_numbers) <- (66) [resolution 308,728] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1004. 95 | ~66 [avatar split clause 1001,726,914]
1118. v1_xxreal_0(k6_numbers) <- (95) [resolution 916,286] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1119. v1_xreal_0(k6_numbers) <- (95) [resolution 916,285] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1120. v1_xcmplx_0(k6_numbers) <- (95) [resolution 916,284] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1124. 123 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1127. 123 | ~95 [avatar split clause 1118,914,1124]
1128. 64 | ~95 [avatar split clause 1119,914,710]
1130. 124 <=> v1_xcmplx_0(k6_numbers) [avatar definition]
1133. 124 | ~95 [avatar split clause 1120,914,1130]
1197. k6_numbers = k6_square_1(k6_numbers) | ~m1_subset_1(k6_numbers,k1_numbers) <- (15) [superposition 469,314] {a:1,w:7,wCS:2800,sine:255,thAx:0,allAx:2,thDist:-2}
1201. 133 <=> m1_subset_1(k6_numbers,k1_numbers) [avatar definition]
1205. 134 <=> k6_numbers = k6_square_1(k6_numbers) [avatar definition]
1209. ~133 | 134 | ~15 [avatar split clause 1197,467,1205,1201]
1225. 135 <=> v1_xboole_0(sK0) [avatar definition]
1227. v1_xboole_0(sK0) <- (135) [avatar component clause 1225] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1229. 136 <=> v2_xxreal_0(sK0) [avatar definition]
1231. v2_xxreal_0(sK0) <- (136) [avatar component clause 1229] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1233. 137 <=> v3_xxreal_0(sK0) [avatar definition]
1234. ~v3_xxreal_0(sK0) <- (~137) [avatar component clause 1233] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1238. 138 <=> v1_xboole_0(sK1) [avatar definition]
1240. v1_xboole_0(sK1) <- (138) [avatar component clause 1238] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1242. 139 <=> v2_xxreal_0(sK1) [avatar definition]
1246. 140 <=> v3_xxreal_0(sK1) [avatar definition]
1247. ~v3_xxreal_0(sK1) <- (~140) [avatar component clause 1246] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1301. $false <- (~2, 12, 135) [unit resulting resolution 403,1227,1000] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1308. 2 | ~12 | ~135 [avatar contradiction clause 1301]
1372. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~108) [resolution 350,998] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1379. ~v1_xreal_0(sK0) <- (6, ~108) [subsumption resolution 1372,424] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1380. $false <- (6, 7, ~108) [subsumption resolution 1379,429] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1381. ~6 | ~7 | 108 [avatar contradiction clause 1380]
1384. $false <- (~3, 12, 138) [unit resulting resolution 407,1240,1000] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1391. 3 | ~12 | ~138 [avatar contradiction clause 1384]
1428. 164 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1429. ~v3_xxreal_0(k6_numbers) <- (~164) [avatar component clause 1428] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1430. v3_xxreal_0(k6_numbers) <- (164) [avatar component clause 1428] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1483. 170 <=> v3_xxreal_0(k2_xcmplx_0(sK0,sK1)) [avatar definition]
1485. v3_xxreal_0(k2_xcmplx_0(sK0,sK1)) <- (170) [avatar component clause 1483] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1512. 172 <=> r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,sK1)) [avatar definition]
1514. ~r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,sK1)) <- (~172) [avatar component clause 1512] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1516. 173 <=> k6_numbers = k2_xcmplx_0(sK0,sK1) [avatar definition]
1518. k6_numbers = k2_xcmplx_0(sK0,sK1) <- (173) [avatar component clause 1516] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1589. 178 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1590. v2_xxreal_0(k6_numbers) <- (178) [avatar component clause 1589] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1591. ~v2_xxreal_0(k6_numbers) <- (~178) [avatar component clause 1589] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1649. m1_subset_1(k6_numbers,k1_numbers) | ~m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) | v1_xboole_0(k4_ordinal1) | v1_xboole_0(k1_numbers) <- (106) [resolution 343,981] {a:1,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
1653. m1_subset_1(k6_numbers,k1_numbers) | v1_xboole_0(k4_ordinal1) | v1_xboole_0(k1_numbers) <- (105, 106) [subsumption resolution 1649,972] {a:1,w:7,wCS:2800,sine:255,thAx:0,allAx:2,thDist:-2}
1654. m1_subset_1(k6_numbers,k1_numbers) | v1_xboole_0(k1_numbers) <- (~20, 105, 106) [subsumption resolution 1653,494] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:2,thDist:-2}
1655. m1_subset_1(k6_numbers,k1_numbers) <- (~8, ~20, 105, 106) [subsumption resolution 1654,434] {a:1,w:3,wCS:1200,sine:255,thAx:0,allAx:2,thDist:-2}
1656. 133 | 8 | 20 | ~105 | ~106 [avatar split clause 1655,979,970,492,432,1201]
1768. v2_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) | ~v1_xreal_0(sK1) | ~v2_xxreal_0(sK1) <- (173) [superposition 355,1518] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1769. v2_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) <- (173) [superposition 356,1518] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1820. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (164) [resolution 323,1430] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1825. ~v1_xxreal_0(k6_numbers) <- (65, 164) [subsumption resolution 1820,722] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1826. ~123 | ~65 | ~164 [avatar split clause 1825,1428,720,1124]
1840. v2_xxreal_0(k6_numbers) | v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) <- (6, 173) [subsumption resolution 1769,424] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1841. v2_xxreal_0(k6_numbers) | v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (6, 7, 173) [subsumption resolution 1840,429] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1842. v2_xxreal_0(k6_numbers) | v3_xxreal_0(sK1) <- (6, 7, 136, 173) [subsumption resolution 1841,1231] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1843. 140 | 178 | ~6 | ~7 | ~136 | ~173 [avatar split clause 1842,1516,1229,427,422,1589,1246]
1844. v2_xxreal_0(k6_numbers) | v3_xxreal_0(sK0) | ~v1_xreal_0(sK1) | ~v2_xxreal_0(sK1) <- (7, 173) [subsumption resolution 1768,429] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1845. v2_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v2_xxreal_0(sK1) <- (7, ~137, 173) [subsumption resolution 1844,1234] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1846. v2_xxreal_0(k6_numbers) | ~v2_xxreal_0(sK1) <- (6, 7, ~137, 173) [subsumption resolution 1845,424] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1847. ~139 | 178 | ~6 | ~7 | 137 | ~173 [avatar split clause 1846,1516,1233,427,422,1589,1242]
1863. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (178) [resolution 1590,326] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
1865. ~v1_xxreal_0(k6_numbers) <- (65, 178) [subsumption resolution 1863,722] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
1866. ~123 | ~65 | ~178 [avatar split clause 1865,1589,720,1124]
1902. 186 <=> k6_numbers = k6_square_1(k2_xcmplx_0(k6_numbers,sK1)) [avatar definition]
1903. k6_numbers != k6_square_1(k2_xcmplx_0(k6_numbers,sK1)) <- (~186) [avatar component clause 1902] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
2177. k6_numbers != k6_square_1(k2_xcmplx_0(k6_numbers,sK1)) <- (~1, 2) [forward demodulation 399,402] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
2178. ~186 | 1 | ~2 [avatar split clause 2177,401,397,1902]
2212. k6_numbers != k6_square_1(k2_xcmplx_0(k6_numbers,k6_numbers)) <- (3, ~186) [forward demodulation 1903,406] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
2214. 213 <=> k6_numbers = k6_square_1(k2_xcmplx_0(k6_numbers,k6_numbers)) [avatar definition]
2216. k6_numbers != k6_square_1(k2_xcmplx_0(k6_numbers,k6_numbers)) <- (~213) [avatar component clause 2214] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
2217. ~213 | ~3 | 186 [avatar split clause 2212,1902,405,2214]
2220. k6_numbers != k6_square_1(k6_numbers) | ~v1_xcmplx_0(k6_numbers) <- (~213) [superposition 2216,290] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2221. ~124 | ~134 | 213 [avatar split clause 2220,2214,1205,1130]
2222. v1_xxreal_0(sK1) <- (6) [resolution 424,292] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
2224. 88 | ~6 [avatar split clause 2222,422,869]
2226. v1_xxreal_0(sK0) <- (7) [resolution 429,292] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
2228. 87 | ~7 [avatar split clause 2226,427,864]
2233. ~v3_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (4) [resolution 414,298] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2239. ~v3_xxreal_0(sK1) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (4, ~164) [subsumption resolution 2233,1429] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2240. ~v3_xxreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (4, 6, ~164) [subsumption resolution 2239,424] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2241. ~v3_xxreal_0(sK1) <- (4, 6, 64, ~164) [subsumption resolution 2240,712] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2242. ~140 | ~4 | ~6 | ~64 | 164 [avatar split clause 2241,1428,710,422,412,1246]
2246. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (5) [resolution 419,298] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2252. ~v3_xxreal_0(sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (5, ~164) [subsumption resolution 2246,1429] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2253. ~v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (5, 7, ~164) [subsumption resolution 2252,429] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2254. ~v3_xxreal_0(sK0) <- (5, 7, 64, ~164) [subsumption resolution 2253,712] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2255. ~137 | ~5 | ~7 | ~64 | 164 [avatar split clause 2254,1428,710,427,417,1233]
2256. v3_xxreal_0(sK0) | v2_xxreal_0(sK0) | v1_xboole_0(sK0) <- (87) [resolution 866,317] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
2257. 135 | 136 | 137 | ~87 [avatar split clause 2256,864,1233,1229,1225]
2258. k6_numbers != k6_numbers | k6_numbers = k2_xcmplx_0(sK0,sK1) | ~r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k2_xcmplx_0(sK0,sK1)) <- (1) [superposition 295,398] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2260. k6_numbers = k2_xcmplx_0(sK0,sK1) | ~r1_xxreal_0(k6_numbers,k2_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k2_xcmplx_0(sK0,sK1)) <- (1) [trivial inequality removal 2258] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
2261. ~108 | ~172 | 173 | ~1 [avatar split clause 2260,397,1516,1512,996]
2262. v3_xxreal_0(sK1) | v2_xxreal_0(sK1) | v1_xboole_0(sK1) <- (88) [resolution 871,317] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
2263. 138 | 139 | 140 | ~88 [avatar split clause 2262,869,1246,1242,1238]
2311. v3_xxreal_0(k2_xcmplx_0(sK0,sK1)) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k2_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (~172) [resolution 1514,300] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2314. v3_xxreal_0(k2_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k2_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (~172, ~178) [subsumption resolution 2311,1591] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2315. v3_xxreal_0(k2_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (108, ~172, ~178) [subsumption resolution 2314,997] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2316. v3_xxreal_0(k2_xcmplx_0(sK0,sK1)) <- (64, 108, ~172, ~178) [subsumption resolution 2315,712] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
2317. 170 | ~64 | ~108 | 172 | 178 [avatar split clause 2316,1589,1512,996,710,1483]
2647. ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (170) [resolution 1485,342] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2651. v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (6, 170) [subsumption resolution 2647,424] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2652. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (6, ~140, 170) [subsumption resolution 2651,1247] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2653. v3_xxreal_0(sK0) <- (6, 7, ~140, 170) [subsumption resolution 2652,429] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2654. $false <- (6, 7, ~137, ~140, 170) [subsumption resolution 2653,1234] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2655. ~6 | ~7 | 137 | 140 | ~170 [avatar contradiction clause 2654]
2657. $false [avatar sat refutation 408,409,410,415,420,425,430,435,440,450,455,465,470,475,480,495,723,729,973,982,1004,1127,1128,1133,1209,1308,1381,1391,1656,1826,1843,1847,1866,2178,2217,2221,2224,2228,2242,2255,2257,2261,2263,2317,2655] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:2410,thDist:-2410}
% SZS output end Proof for square_1__t31_square_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 104
% >>> Preprocessing
% Introduced skolems: 21
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 137
% Generated clauses: 1413
% Active clauses: 310
% Passive clauses: 510
% Final active clauses: 232
% Final passive clauses: 139
% Main loop iterations started: 347
% 
% >>> Simplifying Inferences
% Duplicate literals: 79
% Trivial inequalities: 6
% Fw subsumption resolutions: 193
% Fw demodulations: 29
% Bw demodulations: 28
% 
% >>> Deletion Inferences
% Simple tautologies: 52
% Forward subsumptions: 200
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 430
% Unit resulting resolution: 14
% Factoring: 1
% Forward superposition: 103
% Backward superposition: 61
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 69
% Split components: 158
% Sat splitting refutations: 22
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 381
% SAT solver unit clauses: 63
% SAT solver binary clauses: 140
% 
% Memory used [KB]: 6780
% Time elapsed: 0.013 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% sat proof minimization: 0.001 s
% congruence closure: 0.002 s
% forward subsumption: 0.001 s
% resolution: 0.003 s
% passive container maintenance: 0.001 s
% other: 0.013 s ( own 0.004 s ) 

