i: [1,51,0,2,1,1,0,0] t146_pboole
a: [3,100,0,2,1,51]
s: 100
i: [1,49,0,2,1,1,0,0] t146_pboole
a: [3,110,0,2,1,49]
s: 110
i: [1,47,0,2,1,1,0,0] t146_pboole
a: [3,120,0,2,1,47]
s: 120
i: [1,45,0,2,1,1,0,0] t146_pboole
a: [3,130,0,2,1,45]
s: 130
i: [1,52,0,3,1,1,0,0] t146_pboole
a: [3,95,0,3,1,52]
s: 95
i: [1,50,0,3,1,1,0,0] t146_pboole
a: [3,105,0,3,1,50]
s: 105
i: [1,48,0,3,1,1,0,0] t146_pboole
a: [3,115,0,3,1,48]
s: 115
i: [1,46,0,3,1,1,0,0] t146_pboole
a: [3,125,0,3,1,46]
s: 125
i: [1,55,0,4,1,1,0,0] t146_pboole
a: [3,80,0,4,1,55]
s: 80
i: [1,54,0,4,1,1,0,0] t146_pboole
a: [3,85,0,4,1,54]
s: 85
i: [1,53,0,4,1,1,0,0] t146_pboole
a: [3,90,0,4,1,53]
s: 90
i: [1,72,0,2,1,0,0,1] rc1_funct_1
a: [3,135,0,2,1,72]
s: 135
i: [1,71,0,2,1,0,0,1] rc1_funct_1
a: [3,140,0,2,1,71]
s: 140
i: [1,57,0,3,1,0,0,1] rc1_pboole
s: 57
i: [1,59,0,3,1,0,0,1] rc1_pboole
s: 59
i: [1,58,0,4,1,0,0,1] rc1_pboole
s: 58
i: [1,60,0,4,1,0,0,1] rc1_pboole
s: 60
i: [1,74,0,24,9,0,0,255] redefinition_r6_pboole
d: [2,75,0,14,5,0,18,74]
s: 75
i: [1,73,0,24,9,0,0,255] redefinition_r6_pboole
d: [2,76,0,14,5,0,18,73]
s: 76
i: [1,61,0,24,9,0,0,1] reflexivity_r2_pboole
s: 61
i: [1,62,0,24,9,0,0,1] reflexivity_r2_pboole
s: 62
i: [1,63,0,24,9,0,0,1] reflexivity_r6_pboole
s: 63
i: [1,64,0,24,9,0,0,1] reflexivity_r6_pboole
s: 64
i: [1,67,0,27,10,0,0,255] redefinition_r6_pboole
s: 67
i: [1,69,0,27,10,0,0,255] redefinition_r6_pboole
s: 69
i: [1,65,0,28,10,0,0,1] symmetry_r6_pboole
s: 65
i: [1,66,0,28,10,0,0,1] symmetry_r6_pboole
s: 66
i: [1,56,0,31,11,0,0,255] l22_pboole
s: 56
d: [2,146,1,27,10,1,40,56,85]
d: [2,150,1,23,9,2,21,146,90]
d: [2,151,1,20,8,3,21,150,95]
d: [2,152,1,18,7,4,21,151,100]
d: [2,153,1,15,6,5,21,152,105]
d: [2,154,1,13,5,6,21,153,110]
d: [2,155,1,10,4,7,21,154,115]
d: [2,156,1,8,3,8,21,155,120]
d: [2,157,1,5,2,9,21,156,125]
d: [2,158,1,3,1,10,21,157,130]
a: [3,162,1,3,1,158]
s: 162
r: 80
r: 85
r: 90
r: 95
r: 100
r: 105
r: 110
d: [2,174,0,4,1,2,23,80,162]
a: [3,184,0,4,1,174]
s: 184
d: [2,192,1,0,0,5,52,130,120,125,115,184,75]
e: 192
d: [2,193,1,10,4,1,40,184,75]
d: [2,195,1,7,3,2,21,193,115]
d: [2,196,1,5,2,3,21,195,120]
d: [2,197,1,2,1,4,21,196,125]
d: [2,198,1,0,0,5,21,197,130]
e: 198
d: [2,147,1,27,10,1,40,56,90]
d: [2,164,1,23,9,2,21,147,85]
d: [2,165,1,20,8,3,21,164,115]
d: [2,166,1,18,7,4,21,165,120]
d: [2,167,1,15,6,5,21,166,125]
d: [2,168,1,13,5,6,21,167,130]
d: [2,169,1,10,4,7,21,168,95]
d: [2,170,1,8,3,8,21,169,100]
d: [2,171,1,5,2,9,21,170,105]
d: [2,172,1,3,1,10,21,171,110]
f: 198,174,172,45,46,47,48,49,50,51,52,53,54,55
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pboole__t146_pboole
% SZS output start Proof for pboole__t146_pboole
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => r6_pboole(X0,X1,X2)))) [input t146_pboole]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => r6_pboole(X0,X1,X2)))) [negated conjecture 1]
10. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (r6_pboole(X0,X1,X2) <=> X1 = X2)) [input redefinition_r6_pboole]
13. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => X1 = X2))) [input l22_pboole]
14. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => r6_pboole(X0,X1,X2)))) [flattening 2]
15. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ((r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2)) => X1 = X2))) [flattening 13]
16. ? [X0,X1] : (? [X2] : ((~r6_pboole(X0,X1,X2) & (r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2))) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 14]
17. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,X1,X2) & r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 16]
18. ! [X0,X1] : (! [X2] : ((X1 = X2 | (~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 15]
19. ! [X0,X1] : (! [X2] : (X1 = X2 | ~r2_pboole(X0,X2,X1) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 18]
34. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 10]
35. ! [X0,X1,X2] : ((r6_pboole(X0,X1,X2) <=> X1 = X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 34]
36. ? [X0,X1] : (? [X2] : (~r6_pboole(X0,X1,X2) & r2_pboole(X0,X2,X1) & r2_pboole(X0,X1,X2) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (~r6_pboole(sK0,sK1,X2) & r2_pboole(sK0,X2,sK1) & r2_pboole(sK0,sK1,X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
37. ? [X2] : (~r6_pboole(sK0,sK1,X2) & r2_pboole(sK0,X2,sK1) & r2_pboole(sK0,sK1,X2) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (~r6_pboole(sK0,sK1,sK2) & r2_pboole(sK0,sK2,sK1) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
38. (~r6_pboole(sK0,sK1,sK2) & r2_pboole(sK0,sK2,sK1) & r2_pboole(sK0,sK1,sK2) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 17,37,36]
42. ! [X0,X1,X2] : (((r6_pboole(X0,X1,X2) | X1 != X2) & (X1 = X2 | ~r6_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 35]
45. v1_relat_1(sK1) [cnf transformation 38] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
46. v4_relat_1(sK1,sK0) [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
47. v1_funct_1(sK1) [cnf transformation 38] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. v1_partfun1(sK1,sK0) [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. v1_relat_1(sK2) [cnf transformation 38] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. v4_relat_1(sK2,sK0) [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. v1_funct_1(sK2) [cnf transformation 38] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. v1_partfun1(sK2,sK0) [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
53. r2_pboole(sK0,sK1,sK2) [cnf transformation 38] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. r2_pboole(sK0,sK2,sK1) [cnf transformation 38] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. ~r6_pboole(sK0,sK1,sK2) [cnf transformation 38] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. ~r2_pboole(X0,X2,X1) | X1 = X2 | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 19] {a:0,w:31,wCS:12400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
70. r6_pboole(X0,X1,X2) | X1 != X2 | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 42] {a:0,w:27,wCS:10800,sine:255,thAx:0,allAx:1,thDist:-1}
74. r6_pboole(X0,X2,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) [equality resolution 70] {a:0,w:24,wCS:9600,sine:255,thAx:0,allAx:1,thDist:-1}
75. r6_pboole(X0,X2,X2) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) [duplicate literal removal 74] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
78. 1 <=> r6_pboole(sK0,sK1,sK2) [avatar definition]
80. ~r6_pboole(sK0,sK1,sK2) <- (~1) [avatar component clause 78] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
81. ~1 [avatar split clause 55,78]
83. 2 <=> r2_pboole(sK0,sK2,sK1) [avatar definition]
85. r2_pboole(sK0,sK2,sK1) <- (2) [avatar component clause 83] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
86. 2 [avatar split clause 54,83]
88. 3 <=> r2_pboole(sK0,sK1,sK2) [avatar definition]
90. r2_pboole(sK0,sK1,sK2) <- (3) [avatar component clause 88] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. 3 [avatar split clause 53,88]
93. 4 <=> v1_partfun1(sK2,sK0) [avatar definition]
95. v1_partfun1(sK2,sK0) <- (4) [avatar component clause 93] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. 4 [avatar split clause 52,93]
98. 5 <=> v1_funct_1(sK2) [avatar definition]
100. v1_funct_1(sK2) <- (5) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. 5 [avatar split clause 51,98]
103. 6 <=> v4_relat_1(sK2,sK0) [avatar definition]
105. v4_relat_1(sK2,sK0) <- (6) [avatar component clause 103] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
106. 6 [avatar split clause 50,103]
108. 7 <=> v1_relat_1(sK2) [avatar definition]
110. v1_relat_1(sK2) <- (7) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
111. 7 [avatar split clause 49,108]
113. 8 <=> v1_partfun1(sK1,sK0) [avatar definition]
115. v1_partfun1(sK1,sK0) <- (8) [avatar component clause 113] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. 8 [avatar split clause 48,113]
118. 9 <=> v1_funct_1(sK1) [avatar definition]
120. v1_funct_1(sK1) <- (9) [avatar component clause 118] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
121. 9 [avatar split clause 47,118]
123. 10 <=> v4_relat_1(sK1,sK0) [avatar definition]
125. v4_relat_1(sK1,sK0) <- (10) [avatar component clause 123] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. 10 [avatar split clause 46,123]
128. 11 <=> v1_relat_1(sK1) [avatar definition]
130. v1_relat_1(sK1) <- (11) [avatar component clause 128] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
131. 11 [avatar split clause 45,128]
147. sK1 = sK2 | ~r2_pboole(sK0,sK2,sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (3) [resolution 56,90] {a:1,w:27,wCS:2700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
160. 14 <=> sK1 = sK2 [avatar definition]
162. sK1 = sK2 <- (14) [avatar component clause 160] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
164. sK1 = sK2 | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (2, 3) [subsumption resolution 147,85] {a:1,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
165. sK1 = sK2 | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (2, 3, 8) [subsumption resolution 164,115] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
166. sK1 = sK2 | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (2, 3, 8, 9) [subsumption resolution 165,120] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
167. sK1 = sK2 | ~v1_relat_1(sK1) | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (2, 3, 8, 9, 10) [subsumption resolution 166,125] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
168. sK1 = sK2 | ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (2, 3, 8, 9, 10, 11) [subsumption resolution 167,130] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
169. sK1 = sK2 | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (2, 3, 4, 8, 9, 10, 11) [subsumption resolution 168,95] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
170. sK1 = sK2 | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) <- (2, 3, 4, 5, 8, 9, 10, 11) [subsumption resolution 169,100] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
171. sK1 = sK2 | ~v1_relat_1(sK2) <- (2, 3, 4, 5, 6, 8, 9, 10, 11) [subsumption resolution 170,105] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
172. sK1 = sK2 <- (2, 3, 4, 5, 6, 7, 8, 9, 10, 11) [subsumption resolution 171,110] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
173. 14 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 [avatar split clause 172,128,123,118,113,108,103,98,93,88,83,160]
174. ~r6_pboole(sK0,sK1,sK1) <- (~1, 14) [backward demodulation 80,162] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
182. 15 <=> r6_pboole(sK0,sK1,sK1) [avatar definition]
184. ~r6_pboole(sK0,sK1,sK1) <- (~15) [avatar component clause 182] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
185. ~15 | 1 | ~14 [avatar split clause 174,160,78,182]
193. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~15) [resolution 184,75] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
195. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (8, ~15) [subsumption resolution 193,115] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
196. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (8, 9, ~15) [subsumption resolution 195,120] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
197. ~v1_relat_1(sK1) <- (8, 9, 10, ~15) [subsumption resolution 196,125] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
198. $false <- (8, 9, 10, 11, ~15) [subsumption resolution 197,130] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
199. ~8 | ~9 | ~10 | ~11 | 15 [avatar contradiction clause 198]
200. $false [avatar sat refutation 81,86,91,96,101,106,111,116,121,126,131,173,185,199] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:118,thDist:-118}
% SZS output end Proof for pboole__t146_pboole
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 28
% Generated clauses: 85
% Active clauses: 24
% Passive clauses: 31
% Final active clauses: 17
% Final passive clauses: 1
% Main loop iterations started: 30
% 
% >>> Simplifying Inferences
% Duplicate literals: 32
% Fw subsumption resolutions: 22
% Bw demodulations: 7
% 
% >>> Deletion Inferences
% Equational tautologies: 2
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 6
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 22
% SAT solver unit clauses: 13
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit clause index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

