i: [1,71,0,2,1,1,0,0] t32_wellord1
a: [3,111,0,2,1,71]
s: 111
i: [1,70,0,2,1,1,0,0] t32_wellord1
a: [3,116,0,2,1,70]
s: 116
i: [1,72,0,10,1,1,0,0] t32_wellord1
a: [3,106,0,10,1,72]
s: 106
i: [1,91,0,3,1,0,0,2] reflexivity_r1_tarski
s: 91
i: [1,90,0,4,1,0,0,255] existence_m1_subset_1
s: 90
i: [1,92,0,4,1,0,0,255] fc6_relat_1
s: 92
i: [1,93,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 93
i: [1,94,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 94
i: [1,81,0,6,3,0,0,255] cc1_wellord1
s: 81
i: [1,82,0,6,3,0,0,255] cc1_wellord1
s: 82
i: [1,83,0,6,3,0,0,255] cc1_wellord1
s: 83
d: [2,118,1,4,2,1,40,81,111]
d: [2,119,1,2,1,2,21,118,116]
a: [3,123,1,2,1,119]
s: 123
i: [1,84,0,6,3,0,0,255] cc1_wellord1
s: 84
d: [2,125,1,4,2,1,40,82,111]
d: [2,126,1,2,1,2,21,125,116]
a: [3,130,1,2,1,126]
s: 130
i: [1,85,0,6,3,0,0,255] cc1_wellord1
s: 85
i: [1,95,0,6,1,0,0,255] dt_k6_subset_1
s: 95
d: [2,132,1,4,2,1,40,83,111]
d: [2,133,1,2,1,2,21,132,116]
a: [3,137,1,2,1,133]
s: 137
i: [1,99,0,6,2,0,0,255] fc2_relat_1
s: 99
d: [2,139,1,4,2,1,40,84,111]
d: [2,140,1,2,1,2,21,139,116]
a: [3,144,1,2,1,140]
s: 144
i: [1,100,0,6,2,0,0,1] dt_k2_wellord1
s: 100
i: [1,101,0,6,2,0,0,255] fc1_relat_1
s: 101
d: [2,146,1,4,2,1,40,85,111]
d: [2,147,1,2,1,2,21,146,116]
a: [3,151,1,2,1,147]
s: 151
i: [1,88,0,7,2,0,0,255] t3_subset
s: 88
d: [2,154,1,5,1,0,40,88,90]
s: 154
i: [1,89,0,7,2,0,0,255] t3_subset
s: 89
i: [1,96,0,7,1,0,0,255] redefinition_k6_subset_1
s: 96
r: 99
d: [2,155,1,5,1,0,40,88,95]
s: 155
d: [2,157,0,6,2,0,23,99,96]
s: 157
i: [1,97,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 97
i: [1,98,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 98
i: [1,77,0,8,3,0,0,255] cc2_relat_1
s: 77
d: [2,158,1,6,2,0,44,101,97]
s: 158
i: [1,86,0,8,2,0,0,1] t9_wellord1
s: 86
d: [2,159,1,6,2,0,44,101,97]
s: 159
i: [1,102,0,8,3,0,0,255] fc3_relat_1
s: 102
i: [1,73,0,10,2,0,0,255] d6_relat_1
s: 73
d: [2,160,1,6,2,0,40,77,90]
s: 160
i: [1,74,0,10,2,0,0,255] d17_relat_1
s: 74
d: [2,162,1,7,3,0,40,77,89]
s: 162
i: [1,75,0,11,2,0,0,255] d6_wellord1
s: 75
i: [1,76,0,12,2,0,0,255] d1_wellord1
s: 76
d: [2,175,1,13,2,0,44,75,93]
d: [2,184,1,9,1,0,21,175,92]
s: 184
i: [1,79,0,14,7,0,0,255] cc2_wellord1
s: 79
d: [2,178,1,13,2,0,44,93,75]
d: [2,185,1,9,1,0,21,178,92]
s: 185
i: [1,87,0,14,4,0,0,1] t31_wellord1
s: 87
d: [2,192,1,17,4,1,44,106,87]
d: [2,194,1,10,3,1,19,192]
d: [2,196,1,8,2,2,21,194,111]
d: [2,197,1,6,1,3,21,196,116]
a: [3,201,1,6,1,197]
s: 201
d: [2,204,2,0,0,2,52,116,201,86]
e: 204
d: [2,205,2,2,1,1,40,201,86]
d: [2,207,2,0,0,2,21,205,116]
e: 207
f: 207,197,70,71,72
% Refutation found. Thanks to Tanya!
% SZS status Theorem for wellord1__t32_wellord1
% SZS output start Proof for wellord1__t32_wellord1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v2_wellord1(X1) => k1_wellord1(X1,X0) = k1_relat_1(k2_wellord1(X1,k1_wellord1(X1,X0))))) [input t32_wellord1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => (v2_wellord1(X1) => k1_wellord1(X1,X0) = k1_relat_1(k2_wellord1(X1,k1_wellord1(X1,X0))))) [negated conjecture 1]
36. ! [X0] : ! [X1] : (v1_relat_1(X1) => ((r1_tarski(X0,k1_relat_1(X1)) & v2_wellord1(X1)) => k1_relat_1(k2_wellord1(X1,X0)) = X0)) [input t31_wellord1]
38. ! [X0] : ! [X1] : (v1_relat_1(X1) => r1_tarski(k1_wellord1(X1,X0),k1_relat_1(X1))) [input t9_wellord1]
39. ~! [X0,X1] : (v1_relat_1(X1) => (v2_wellord1(X1) => k1_wellord1(X1,X0) = k1_relat_1(k2_wellord1(X1,k1_wellord1(X1,X0))))) [flattening 2]
40. ! [X0,X1] : (v1_relat_1(X1) => r1_tarski(k1_wellord1(X1,X0),k1_relat_1(X1))) [flattening 38]
41. ! [X0,X1] : (v1_relat_1(X1) => ((r1_tarski(X0,k1_relat_1(X1)) & v2_wellord1(X1)) => k1_relat_1(k2_wellord1(X1,X0)) = X0)) [flattening 36]
46. ? [X0,X1] : ((k1_wellord1(X1,X0) != k1_relat_1(k2_wellord1(X1,k1_wellord1(X1,X0))) & v2_wellord1(X1)) & v1_relat_1(X1)) [ennf transformation 39]
47. ? [X0,X1] : (k1_wellord1(X1,X0) != k1_relat_1(k2_wellord1(X1,k1_wellord1(X1,X0))) & v2_wellord1(X1) & v1_relat_1(X1)) [flattening 46]
57. ! [X0,X1] : (r1_tarski(k1_wellord1(X1,X0),k1_relat_1(X1)) | ~v1_relat_1(X1)) [ennf transformation 40]
58. ! [X0,X1] : ((k1_relat_1(k2_wellord1(X1,X0)) = X0 | (~r1_tarski(X0,k1_relat_1(X1)) | ~v2_wellord1(X1))) | ~v1_relat_1(X1)) [ennf transformation 41]
59. ! [X0,X1] : (k1_relat_1(k2_wellord1(X1,X0)) = X0 | ~r1_tarski(X0,k1_relat_1(X1)) | ~v2_wellord1(X1) | ~v1_relat_1(X1)) [flattening 58]
65. ? [X0,X1] : (k1_wellord1(X1,X0) != k1_relat_1(k2_wellord1(X1,k1_wellord1(X1,X0))) & v2_wellord1(X1) & v1_relat_1(X1)) => (k1_wellord1(sK1,sK0) != k1_relat_1(k2_wellord1(sK1,k1_wellord1(sK1,sK0))) & v2_wellord1(sK1) & v1_relat_1(sK1)) [choice axiom]
66. k1_wellord1(sK1,sK0) != k1_relat_1(k2_wellord1(sK1,k1_wellord1(sK1,sK0))) & v2_wellord1(sK1) & v1_relat_1(sK1) [skolemisation 47,65]
70. v1_relat_1(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
71. v2_wellord1(sK1) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
72. k1_wellord1(sK1,sK0) != k1_relat_1(k2_wellord1(sK1,k1_wellord1(sK1,sK0))) [cnf transformation 66] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
86. r1_tarski(k1_wellord1(X1,X0),k1_relat_1(X1)) | ~v1_relat_1(X1) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
87. k1_relat_1(k2_wellord1(X1,X0)) = X0 | ~r1_tarski(X0,k1_relat_1(X1)) | ~v2_wellord1(X1) | ~v1_relat_1(X1) [cnf transformation 59] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. 1 <=> k1_wellord1(sK1,sK0) = k1_relat_1(k2_wellord1(sK1,k1_wellord1(sK1,sK0))) [avatar definition]
106. k1_wellord1(sK1,sK0) != k1_relat_1(k2_wellord1(sK1,k1_wellord1(sK1,sK0))) <- (~1) [avatar component clause 104] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
107. ~1 [avatar split clause 72,104]
109. 2 <=> v2_wellord1(sK1) [avatar definition]
111. v2_wellord1(sK1) <- (2) [avatar component clause 109] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
112. 2 [avatar split clause 71,109]
114. 3 <=> v1_relat_1(sK1) [avatar definition]
116. v1_relat_1(sK1) <- (3) [avatar component clause 114] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
117. 3 [avatar split clause 70,114]
192. k1_wellord1(sK1,sK0) != k1_wellord1(sK1,sK0) | ~r1_tarski(k1_wellord1(sK1,sK0),k1_relat_1(sK1)) | ~v2_wellord1(sK1) | ~v1_relat_1(sK1) <- (~1) [superposition 106,87] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
194. ~r1_tarski(k1_wellord1(sK1,sK0),k1_relat_1(sK1)) | ~v2_wellord1(sK1) | ~v1_relat_1(sK1) <- (~1) [trivial inequality removal 192] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
196. ~r1_tarski(k1_wellord1(sK1,sK0),k1_relat_1(sK1)) | ~v1_relat_1(sK1) <- (~1, 2) [subsumption resolution 194,111] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
197. ~r1_tarski(k1_wellord1(sK1,sK0),k1_relat_1(sK1)) <- (~1, 2, 3) [subsumption resolution 196,116] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
199. 9 <=> r1_tarski(k1_wellord1(sK1,sK0),k1_relat_1(sK1)) [avatar definition]
201. ~r1_tarski(k1_wellord1(sK1,sK0),k1_relat_1(sK1)) <- (~9) [avatar component clause 199] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
202. ~9 | 1 | ~2 | ~3 [avatar split clause 197,114,109,104,199]
205. ~v1_relat_1(sK1) <- (~9) [resolution 201,86] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
207. $false <- (3, ~9) [subsumption resolution 205,116] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
208. ~3 | 9 [avatar contradiction clause 207]
209. $false [avatar sat refutation 107,112,117,202,208] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:56,thDist:-56}
% SZS output end Proof for wellord1__t32_wellord1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 37
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 33
% Generated clauses: 101
% Active clauses: 44
% Passive clauses: 57
% Final active clauses: 43
% Final passive clauses: 11
% Main loop iterations started: 46
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Trivial inequalities: 1
% Fw subsumption resolutions: 12
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 17
% Unit resulting resolution: 1
% Forward superposition: 11
% Backward superposition: 14
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 3
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% sine selection: 0.001 s
% term sharing: 0.001 s
% superposition: 0.001 s ( own 0.0000 s ) 
% other: 0.003 s ( own 0.0000 s ) 

