i: [1,74,0,2,1,1,0,0] t60_lukasi_1
a: [3,118,0,2,1,74]
s: 118
i: [1,76,0,5,1,1,0,0] t60_lukasi_1
a: [3,108,0,5,1,76]
s: 108
i: [1,75,0,6,1,1,0,0] t60_lukasi_1
a: [3,113,0,6,1,75]
s: 113
i: [1,77,0,7,1,1,0,0] t60_lukasi_1
a: [3,103,0,7,1,77]
s: 103
i: [1,99,0,2,1,0,0,1] existence_m1_qc_lang1
a: [3,123,0,2,1,99]
s: 123
i: [1,82,0,4,1,0,0,1] existence_m1_subset_1
s: 82
i: [1,78,0,5,2,0,0,255] dt_k9_qc_lang1
s: 78
i: [1,79,0,5,2,0,0,255] fc2_cqc_lang
s: 79
i: [1,80,0,8,2,0,0,1] dt_k3_cqc_lang
s: 80
i: [1,87,0,12,3,0,0,2] fc2_lukasi_1
s: 87
i: [1,86,0,14,4,0,0,255] existence_m2_subset_1
s: 86
i: [1,83,0,15,5,0,0,255] dt_m2_subset_1
s: 83
i: [1,84,0,15,5,0,0,255] redefinition_m2_subset_1
s: 84
d: [2,125,1,14,4,1,40,83,108]
a: [3,135,1,2,1,125]
s: 135
i: [1,85,0,15,5,0,0,255] redefinition_m2_subset_1
s: 85
i: [1,88,0,17,4,0,0,1] dt_k2_qc_lang2
s: 88
d: [2,127,1,18,5,0,40,83,80]
d: [2,140,1,15,4,0,21,127,79]
d: [2,141,1,12,3,0,21,140,78]
s: 141
i: [1,89,0,19,4,0,0,255] redefinition_k8_cqc_lang
s: 89
d: [2,156,2,6,2,1,40,141,113]
d: [2,158,2,4,1,2,21,156,118]
a: [3,162,2,4,1,158]
s: 162
i: [1,90,0,19,4,0,0,2] fc3_lukasi_1
s: 90
i: [1,91,0,19,4,0,0,1] dt_k8_cqc_lang
s: 91
d: [2,157,2,22,5,0,40,141,86]
d: [2,164,2,16,4,0,21,157,80]
d: [2,165,2,13,3,0,21,164,79]
d: [2,166,2,10,2,0,21,165,78]
s: 166
i: [1,81,0,20,5,0,0,1] t59_cqc_the1
s: 81
d: [2,175,1,14,4,1,40,81,113]
d: [2,179,1,12,3,2,21,175,118]
a: [3,186,1,3,1,179]
s: 186
i: [1,92,0,22,4,0,0,2] fc5_lukasi_1
s: 92
i: [1,93,0,25,4,0,0,2] fc6_lukasi_1
s: 93
d: [2,144,1,18,5,0,40,84,80]
d: [2,149,1,15,4,0,21,144,79]
d: [2,150,1,12,3,0,21,149,78]
s: 150
i: [1,94,0,29,5,0,0,2] fc8_lukasi_1
s: 94
d: [2,195,2,6,2,1,40,150,113]
d: [2,199,2,4,1,2,21,195,118]
a: [3,203,2,4,1,199]
s: 203
i: [1,95,0,32,5,0,0,2] fc9_lukasi_1
s: 95
i: [1,96,0,32,5,0,0,2] fc1_lukasi_1
s: 96
d: [2,196,2,22,5,0,40,150,86]
d: [2,205,2,16,4,0,21,196,80]
d: [2,206,2,13,3,0,21,205,79]
d: [2,207,2,10,2,0,21,206,78]
s: 207
i: [1,97,0,32,5,0,0,2] fc4_lukasi_1
s: 97
d: [2,153,1,18,5,0,40,85,80]
d: [2,154,1,15,4,0,21,153,79]
d: [2,155,1,12,3,0,21,154,78]
s: 155
i: [1,98,0,35,5,0,0,2] fc7_lukasi_1
s: 98
d: [2,167,1,22,6,0,44,87,89]
d: [2,170,1,12,3,0,18,167]
s: 170
d: [2,218,2,12,4,0,40,155,141]
d: [2,219,2,10,3,0,18,218]
s: 219
d: [2,126,1,15,4,0,40,83,82]
s: 126
d: [2,225,3,8,2,0,40,219,82]
s: 225
d: [2,143,1,17,4,0,40,84,82]
s: 143
d: [2,152,1,17,4,0,40,85,82]
s: 152
d: [2,228,2,25,6,0,40,126,86]
d: [2,229,2,19,4,0,18,228]
d: [2,230,2,13,3,0,21,229,82]
s: 230
d: [2,176,1,34,7,0,40,81,86]
d: [2,188,1,28,6,0,21,176,80]
d: [2,189,1,25,5,0,21,188,79]
d: [2,190,1,22,4,0,21,189,78]
s: 190
d: [2,235,2,20,6,0,40,152,126]
d: [2,236,2,14,4,0,18,235]
s: 236
d: [2,168,1,27,7,0,44,88,89]
d: [2,169,1,25,6,0,18,168]
s: 169
d: [2,246,1,21,5,0,21,169,219]
d: [2,247,1,17,4,0,21,246,219]
s: 247
d: [2,243,3,12,3,0,40,236,82]
s: 243
d: [2,171,1,32,7,0,44,90,89]
d: [2,172,1,26,5,0,18,171]
s: 172
d: [2,231,2,27,6,0,40,143,86]
d: [2,232,2,21,4,0,18,231]
d: [2,233,2,15,3,0,21,232,82]
s: 233
d: [2,177,1,30,7,0,40,81,91]
d: [2,178,1,28,6,0,18,177]
s: 178
d: [2,191,1,38,7,0,44,92,89]
d: [2,192,1,32,5,0,18,191]
s: 192
d: [2,257,2,21,5,1,40,178,103]
d: [2,258,2,17,4,1,18,257]
d: [2,259,2,12,3,2,21,258,108]
d: [2,260,2,10,2,3,21,259,118]
d: [2,261,2,6,1,4,21,260,203]
a: [3,265,2,6,1,261]
s: 265
d: [2,267,3,0,0,3,52,118,203,265,170]
e: 267
d: [2,268,3,6,2,1,40,265,170]
d: [2,270,3,2,1,2,21,268,203]
d: [2,271,3,0,0,3,21,270,118]
e: 271
f: 271,261,199,74,75,76,77
% Refutation found. Thanks to Tanya!
% SZS status Theorem for lukasi_1__t60_lukasi_1
% SZS output start Proof for lukasi_1__t60_lukasi_1
1. ! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k3_cqc_lang(X0))) => r3_cqc_the1(X0,X2,k8_cqc_lang(X0,X1,X1))))) [input t60_lukasi_1]
2. ~! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k3_cqc_lang(X0))) => r3_cqc_the1(X0,X2,k8_cqc_lang(X0,X1,X1))))) [negated conjecture 1]
5. ! [X0,X1] : ((m1_subset_1(X1,k3_cqc_lang(X0)) & m1_qc_lang1(X0)) => v2_cqc_the1(k2_qc_lang2(X0,X1,X1),X0)) [input fc2_lukasi_1]
16. ! [X0,X1,X2] : ((m1_subset_1(X2,k3_cqc_lang(X0)) & m1_subset_1(X1,k3_cqc_lang(X0)) & m1_qc_lang1(X0)) => k2_qc_lang2(X0,X1,X2) = k8_cqc_lang(X0,X1,X2)) [input redefinition_k8_cqc_lang]
17. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ! [X2] : (m2_subset_1(X2,X0,X1) <=> m1_subset_1(X2,X1))) [input redefinition_m2_subset_1]
19. ! [X0] : (m1_qc_lang1(X0) => m1_subset_1(k3_cqc_lang(X0),k1_zfmisc_1(k9_qc_lang1(X0)))) [input dt_k3_cqc_lang]
20. ! [X0,X1,X2] : ((m1_subset_1(X2,k3_cqc_lang(X0)) & m1_subset_1(X1,k3_cqc_lang(X0)) & m1_qc_lang1(X0)) => m2_subset_1(k8_cqc_lang(X0,X1,X2),k9_qc_lang1(X0),k3_cqc_lang(X0))) [input dt_k8_cqc_lang]
21. ! [X0] : (m1_qc_lang1(X0) => ~v1_xboole_0(k9_qc_lang1(X0))) [input dt_k9_qc_lang1]
25. ! [X0] : (m1_qc_lang1(X0) => ~v1_xboole_0(k3_cqc_lang(X0))) [input fc2_cqc_lang]
26. ! [X0] : (m1_qc_lang1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(k3_cqc_lang(X0))) => ! [X2] : (m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0)) => (v2_cqc_the1(X2,X0) => r3_cqc_the1(X0,X1,X2))))) [input t59_cqc_the1]
27. ? [X0] : (? [X1] : (? [X2] : (~r3_cqc_the1(X0,X2,k8_cqc_lang(X0,X1,X1)) & m1_subset_1(X2,k1_zfmisc_1(k3_cqc_lang(X0)))) & m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0))) & m1_qc_lang1(X0)) [ennf transformation 2]
28. ! [X0] : (~v1_xboole_0(k9_qc_lang1(X0)) | ~m1_qc_lang1(X0)) [ennf transformation 21]
29. ! [X0] : (~v1_xboole_0(k3_cqc_lang(X0)) | ~m1_qc_lang1(X0)) [ennf transformation 25]
30. ! [X0] : (m1_subset_1(k3_cqc_lang(X0),k1_zfmisc_1(k9_qc_lang1(X0))) | ~m1_qc_lang1(X0)) [ennf transformation 19]
31. ! [X0] : (! [X1] : (! [X2] : ((r3_cqc_the1(X0,X1,X2) | ~v2_cqc_the1(X2,X0)) | ~m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(k3_cqc_lang(X0)))) | ~m1_qc_lang1(X0)) [ennf transformation 26]
32. ! [X0] : (! [X1] : (! [X2] : (r3_cqc_the1(X0,X1,X2) | ~v2_cqc_the1(X2,X0) | ~m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(k3_cqc_lang(X0)))) | ~m1_qc_lang1(X0)) [flattening 31]
35. ! [X0,X1] : (! [X2] : (m2_subset_1(X2,X0,X1) <=> m1_subset_1(X2,X1)) | (~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 17]
36. ! [X0,X1] : (! [X2] : (m2_subset_1(X2,X0,X1) <=> m1_subset_1(X2,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 35]
39. ! [X0,X1] : (v2_cqc_the1(k2_qc_lang2(X0,X1,X1),X0) | (~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0))) [ennf transformation 5]
40. ! [X0,X1] : (v2_cqc_the1(k2_qc_lang2(X0,X1,X1),X0) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0)) [flattening 39]
43. ! [X0,X1,X2] : (k2_qc_lang2(X0,X1,X2) = k8_cqc_lang(X0,X1,X2) | (~m1_subset_1(X2,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0))) [ennf transformation 16]
44. ! [X0,X1,X2] : (k2_qc_lang2(X0,X1,X2) = k8_cqc_lang(X0,X1,X2) | ~m1_subset_1(X2,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0)) [flattening 43]
47. ! [X0,X1,X2] : (m2_subset_1(k8_cqc_lang(X0,X1,X2),k9_qc_lang1(X0),k3_cqc_lang(X0)) | (~m1_subset_1(X2,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0))) [ennf transformation 20]
48. ! [X0,X1,X2] : (m2_subset_1(k8_cqc_lang(X0,X1,X2),k9_qc_lang1(X0),k3_cqc_lang(X0)) | ~m1_subset_1(X2,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0)) [flattening 47]
63. ? [X0] : (? [X1] : (? [X2] : (~r3_cqc_the1(X0,X2,k8_cqc_lang(X0,X1,X1)) & m1_subset_1(X2,k1_zfmisc_1(k3_cqc_lang(X0)))) & m2_subset_1(X1,k9_qc_lang1(X0),k3_cqc_lang(X0))) & m1_qc_lang1(X0)) => (? [X1] : (? [X2] : (~r3_cqc_the1(sK0,X2,k8_cqc_lang(sK0,X1,X1)) & m1_subset_1(X2,k1_zfmisc_1(k3_cqc_lang(sK0)))) & m2_subset_1(X1,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) & m1_qc_lang1(sK0)) [choice axiom]
64. ? [X1] : (? [X2] : (~r3_cqc_the1(sK0,X2,k8_cqc_lang(sK0,X1,X1)) & m1_subset_1(X2,k1_zfmisc_1(k3_cqc_lang(sK0)))) & m2_subset_1(X1,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) => (? [X2] : (~r3_cqc_the1(sK0,X2,k8_cqc_lang(sK0,sK1,sK1)) & m1_subset_1(X2,k1_zfmisc_1(k3_cqc_lang(sK0)))) & m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) [choice axiom]
65. ? [X2] : (~r3_cqc_the1(sK0,X2,k8_cqc_lang(sK0,sK1,sK1)) & m1_subset_1(X2,k1_zfmisc_1(k3_cqc_lang(sK0)))) => (~r3_cqc_the1(sK0,sK2,k8_cqc_lang(sK0,sK1,sK1)) & m1_subset_1(sK2,k1_zfmisc_1(k3_cqc_lang(sK0)))) [choice axiom]
66. ((~r3_cqc_the1(sK0,sK2,k8_cqc_lang(sK0,sK1,sK1)) & m1_subset_1(sK2,k1_zfmisc_1(k3_cqc_lang(sK0)))) & m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0))) & m1_qc_lang1(sK0) [skolemisation 27,65,64,63]
69. ! [X0,X1] : (! [X2] : ((m2_subset_1(X2,X0,X1) | ~m1_subset_1(X2,X1)) & (m1_subset_1(X2,X1) | ~m2_subset_1(X2,X0,X1))) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [nnf transformation 36]
74. m1_qc_lang1(sK0) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
75. m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) [cnf transformation 66] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
76. m1_subset_1(sK2,k1_zfmisc_1(k3_cqc_lang(sK0))) [cnf transformation 66] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
77. ~r3_cqc_the1(sK0,sK2,k8_cqc_lang(sK0,sK1,sK1)) [cnf transformation 66] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
78. ~v1_xboole_0(k9_qc_lang1(X0)) | ~m1_qc_lang1(X0) [cnf transformation 28] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
79. ~v1_xboole_0(k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [cnf transformation 29] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
80. m1_subset_1(k3_cqc_lang(X0),k1_zfmisc_1(k9_qc_lang1(X0))) | ~m1_qc_lang1(X0) [cnf transformation 30] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
81. ~m2_subset_1(X2,k9_qc_lang1(X0),k3_cqc_lang(X0)) | ~v2_cqc_the1(X2,X0) | r3_cqc_the1(X0,X1,X2) | ~m1_subset_1(X1,k1_zfmisc_1(k3_cqc_lang(X0))) | ~m1_qc_lang1(X0) [cnf transformation 32] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
84. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | ~m2_subset_1(X2,X0,X1) | m1_subset_1(X2,X1) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 69] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
87. v2_cqc_the1(k2_qc_lang2(X0,X1,X1),X0) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [cnf transformation 40] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
89. k2_qc_lang2(X0,X1,X2) = k8_cqc_lang(X0,X1,X2) | ~m1_subset_1(X2,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [cnf transformation 44] {a:0,w:19,wCS:7600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
91. m2_subset_1(k8_cqc_lang(X0,X1,X2),k9_qc_lang1(X0),k3_cqc_lang(X0)) | ~m1_subset_1(X2,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [cnf transformation 48] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. 1 <=> r3_cqc_the1(sK0,sK2,k8_cqc_lang(sK0,sK1,sK1)) [avatar definition]
103. ~r3_cqc_the1(sK0,sK2,k8_cqc_lang(sK0,sK1,sK1)) <- (~1) [avatar component clause 101] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
104. ~1 [avatar split clause 77,101]
106. 2 <=> m1_subset_1(sK2,k1_zfmisc_1(k3_cqc_lang(sK0))) [avatar definition]
108. m1_subset_1(sK2,k1_zfmisc_1(k3_cqc_lang(sK0))) <- (2) [avatar component clause 106] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
109. 2 [avatar split clause 76,106]
111. 3 <=> m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) [avatar definition]
113. m2_subset_1(sK1,k9_qc_lang1(sK0),k3_cqc_lang(sK0)) <- (3) [avatar component clause 111] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
114. 3 [avatar split clause 75,111]
116. 4 <=> m1_qc_lang1(sK0) [avatar definition]
118. m1_qc_lang1(sK0) <- (4) [avatar component clause 116] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
119. 4 [avatar split clause 74,116]
144. ~m2_subset_1(X3,k9_qc_lang1(X4),k3_cqc_lang(X4)) | m1_subset_1(X3,k3_cqc_lang(X4)) | v1_xboole_0(k3_cqc_lang(X4)) | v1_xboole_0(k9_qc_lang1(X4)) | ~m1_qc_lang1(X4) [resolution 84,80] {a:1,w:18,wCS:7200,sine:1,thAx:0,allAx:2,thDist:-2}
149. ~m2_subset_1(X3,k9_qc_lang1(X4),k3_cqc_lang(X4)) | m1_subset_1(X3,k3_cqc_lang(X4)) | v1_xboole_0(k9_qc_lang1(X4)) | ~m1_qc_lang1(X4) [subsumption resolution 144,79] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
150. ~m2_subset_1(X3,k9_qc_lang1(X4),k3_cqc_lang(X4)) | m1_subset_1(X3,k3_cqc_lang(X4)) | ~m1_qc_lang1(X4) [subsumption resolution 149,78] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
167. v2_cqc_the1(k8_cqc_lang(X0,X1,X1),X0) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [superposition 87,89] {a:1,w:22,wCS:8800,sine:2,thAx:0,allAx:2,thDist:-2}
170. v2_cqc_the1(k8_cqc_lang(X0,X1,X1),X0) | ~m1_subset_1(X1,k3_cqc_lang(X0)) | ~m1_qc_lang1(X0) [duplicate literal removal 167] {a:1,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
177. ~v2_cqc_the1(k8_cqc_lang(X3,X4,X5),X3) | r3_cqc_the1(X3,X6,k8_cqc_lang(X3,X4,X5)) | ~m1_subset_1(X6,k1_zfmisc_1(k3_cqc_lang(X3))) | ~m1_qc_lang1(X3) | ~m1_subset_1(X5,k3_cqc_lang(X3)) | ~m1_subset_1(X4,k3_cqc_lang(X3)) | ~m1_qc_lang1(X3) [resolution 81,91] {a:1,w:30,wCS:12000,sine:1,thAx:0,allAx:2,thDist:-2}
178. r3_cqc_the1(X3,X6,k8_cqc_lang(X3,X4,X5)) | ~v2_cqc_the1(k8_cqc_lang(X3,X4,X5),X3) | ~m1_subset_1(X6,k1_zfmisc_1(k3_cqc_lang(X3))) | ~m1_qc_lang1(X3) | ~m1_subset_1(X5,k3_cqc_lang(X3)) | ~m1_subset_1(X4,k3_cqc_lang(X3)) [duplicate literal removal 177] {a:1,w:28,wCS:11200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
195. m1_subset_1(sK1,k3_cqc_lang(sK0)) | ~m1_qc_lang1(sK0) <- (3) [resolution 150,113] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
199. m1_subset_1(sK1,k3_cqc_lang(sK0)) <- (3, 4) [subsumption resolution 195,118] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
201. 13 <=> m1_subset_1(sK1,k3_cqc_lang(sK0)) [avatar definition]
203. m1_subset_1(sK1,k3_cqc_lang(sK0)) <- (13) [avatar component clause 201] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
204. 13 | ~3 | ~4 [avatar split clause 199,116,111,201]
257. ~v2_cqc_the1(k8_cqc_lang(sK0,sK1,sK1),sK0) | ~m1_subset_1(sK2,k1_zfmisc_1(k3_cqc_lang(sK0))) | ~m1_qc_lang1(sK0) | ~m1_subset_1(sK1,k3_cqc_lang(sK0)) | ~m1_subset_1(sK1,k3_cqc_lang(sK0)) <- (~1) [resolution 178,103] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
258. ~v2_cqc_the1(k8_cqc_lang(sK0,sK1,sK1),sK0) | ~m1_subset_1(sK2,k1_zfmisc_1(k3_cqc_lang(sK0))) | ~m1_qc_lang1(sK0) | ~m1_subset_1(sK1,k3_cqc_lang(sK0)) <- (~1) [duplicate literal removal 257] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
259. ~v2_cqc_the1(k8_cqc_lang(sK0,sK1,sK1),sK0) | ~m1_qc_lang1(sK0) | ~m1_subset_1(sK1,k3_cqc_lang(sK0)) <- (~1, 2) [subsumption resolution 258,108] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
260. ~v2_cqc_the1(k8_cqc_lang(sK0,sK1,sK1),sK0) | ~m1_subset_1(sK1,k3_cqc_lang(sK0)) <- (~1, 2, 4) [subsumption resolution 259,118] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
261. ~v2_cqc_the1(k8_cqc_lang(sK0,sK1,sK1),sK0) <- (~1, 2, 4, 13) [subsumption resolution 260,203] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
263. 14 <=> v2_cqc_the1(k8_cqc_lang(sK0,sK1,sK1),sK0) [avatar definition]
265. ~v2_cqc_the1(k8_cqc_lang(sK0,sK1,sK1),sK0) <- (~14) [avatar component clause 263] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
266. ~14 | 1 | ~2 | ~4 | ~13 [avatar split clause 261,201,116,106,101,263]
268. ~m1_subset_1(sK1,k3_cqc_lang(sK0)) | ~m1_qc_lang1(sK0) <- (~14) [resolution 265,170] {a:3,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
270. ~m1_qc_lang1(sK0) <- (13, ~14) [subsumption resolution 268,203] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
271. $false <- (4, 13, ~14) [subsumption resolution 270,118] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
272. ~4 | ~13 | 14 [avatar contradiction clause 271]
273. $false [avatar sat refutation 104,109,114,119,204,266,272] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:105,thDist:-105}
% SZS output end Proof for lukasi_1__t60_lukasi_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 143
% Active clauses: 51
% Passive clauses: 71
% Final active clauses: 51
% Final passive clauses: 19
% Main loop iterations started: 52
% 
% >>> Simplifying Inferences
% Duplicate literals: 34
% Fw subsumption resolutions: 29
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 42
% Unit resulting resolution: 1
% Forward superposition: 8
% Backward superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 8
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% backward superposition index maintenance: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

