i: [1,155,0,2,1,1,0,0] t80_partfun1
a: [3,253,0,2,1,155]
s: 253
i: [1,153,0,2,1,1,0,0] t80_partfun1
a: [3,263,0,2,1,153]
s: 263
i: [1,154,0,3,1,1,0,0] t80_partfun1
a: [3,258,0,3,1,154]
s: 258
i: [1,156,0,6,1,1,0,0] t80_partfun1
a: [3,248,0,6,1,156]
s: 248
i: [1,158,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,268,0,2,1,158]
s: 268
i: [1,159,0,2,1,0,0,255] fc1_xboole_0
a: [3,273,0,2,1,159]
s: 273
i: [1,218,0,2,1,0,0,2] rc2_xboole_0
a: [3,283,0,2,1,218]
s: 283
i: [1,219,0,2,1,0,0,2] rc1_xboole_0
a: [3,288,0,2,1,219]
s: 288
i: [1,221,0,2,1,0,0,255] rc7_funct_1
a: [3,293,0,2,1,221]
s: 293
i: [1,220,0,2,1,0,0,255] rc7_funct_1
a: [3,298,0,2,1,220]
s: 298
i: [1,223,0,2,1,0,0,2] rc1_relat_1
a: [3,303,0,2,1,223]
s: 303
i: [1,222,0,2,1,0,0,2] rc1_relat_1
a: [3,308,0,2,1,222]
s: 308
i: [1,227,0,2,1,0,0,3] rc4_funct_1
a: [3,313,0,2,1,227]
s: 313
i: [1,226,0,2,1,0,0,3] rc4_funct_1
a: [3,318,0,2,1,226]
s: 318
i: [1,225,0,2,1,0,0,3] rc4_funct_1
a: [3,323,0,2,1,225]
s: 323
i: [1,224,0,2,1,0,0,3] rc4_funct_1
a: [3,328,0,2,1,224]
s: 328
i: [1,229,0,2,1,0,0,3] rc2_relat_1
a: [3,333,0,2,1,229]
s: 333
i: [1,228,0,2,1,0,0,3] rc2_relat_1
a: [3,338,0,2,1,228]
s: 338
i: [1,232,0,2,1,0,0,3] rc3_funct_1
a: [3,343,0,2,1,232]
s: 343
i: [1,231,0,2,1,0,0,3] rc3_funct_1
a: [3,348,0,2,1,231]
s: 348
i: [1,230,0,2,1,0,0,3] rc3_funct_1
a: [3,353,0,2,1,230]
s: 353
i: [1,234,0,2,1,0,0,1] rc1_funct_1
a: [3,358,0,2,1,234]
s: 358
i: [1,233,0,2,1,0,0,1] rc1_funct_1
a: [3,363,0,2,1,233]
s: 363
i: [1,237,0,2,1,0,0,255] rc5_funct_1
a: [3,368,0,2,1,237]
s: 368
i: [1,236,0,2,1,0,0,255] rc5_funct_1
a: [3,373,0,2,1,236]
s: 373
i: [1,235,0,2,1,0,0,255] rc5_funct_1
a: [3,378,0,2,1,235]
s: 378
i: [1,239,0,2,1,0,0,255] rc2_funct_1
a: [3,383,0,2,1,239]
s: 383
i: [1,238,0,2,1,0,0,255] rc2_funct_1
a: [3,388,0,2,1,238]
s: 388
i: [1,157,0,11,1,1,0,0] t80_partfun1
a: [3,243,0,11,1,157]
s: 243
i: [1,160,0,3,1,0,0,255] d2_xboole_0
a: [3,278,0,3,1,160]
s: 278
r: 268
i: [1,163,0,4,2,0,0,255] cc7_funct_1
s: 163
d: [2,391,1,2,1,1,40,163,288]
a: [3,396,1,2,1,391]
s: 396
i: [1,164,0,4,2,0,0,1] cc1_funct_1
s: 164
d: [2,392,1,2,1,1,40,163,273]
a: [3,401,1,2,1,392]
s: 401
i: [1,165,0,4,2,0,0,2] cc1_relat_1
s: 165
i: [1,199,0,4,1,0,0,2] existence_m1_subset_1
s: 199
d: [2,403,1,2,1,1,40,164,288]
a: [3,408,1,2,1,403]
s: 408
i: [1,166,0,5,2,0,0,255] t6_boole
s: 166
d: [2,404,1,2,1,1,40,164,273]
a: [3,413,1,2,1,404]
s: 413
i: [1,167,0,5,2,0,0,1] fc10_relat_1
s: 167
i: [1,194,0,5,2,0,0,1] t7_boole
s: 194
d: [2,434,1,5,1,1,40,194,248]
a: [3,438,1,5,1,434]
s: 438
i: [1,184,0,6,3,0,0,255] cc3_relat_1
s: 184
d: [2,440,2,4,1,1,40,438,167]
a: [3,444,2,4,1,440]
s: 444
d: [2,446,0,4,2,0,21,184,165]
s: 446
i: [1,186,0,6,3,0,0,3] cc4_relat_1
s: 186
d: [2,415,1,2,1,1,40,165,288]
a: [3,420,1,2,1,415]
s: 420
d: [2,460,0,4,2,0,21,186,165]
s: 460
d: [2,416,1,2,1,1,40,165,273]
a: [3,425,1,2,1,416]
s: 425
i: [1,190,0,6,2,0,0,2] t1_subset
s: 190
i: [1,200,0,6,2,0,0,1] dt_k5_relat_1
s: 200
d: [2,475,1,6,1,1,40,190,248]
a: [3,479,1,6,1,475]
s: 479
i: [1,202,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 202
d: [2,481,1,6,1,1,40,202,248]
a: [3,485,1,6,1,481]
s: 485
i: [1,161,0,7,3,0,0,255] cc8_funct_1
s: 161
i: [1,162,0,7,3,0,0,255] cc8_funct_1
s: 162
d: [2,488,2,7,2,1,40,161,479]
a: [3,492,2,5,1,488]
s: 492
i: [1,173,0,7,3,0,0,1] fc8_relat_1
s: 173
d: [2,447,1,2,1,1,40,446,288]
a: [3,453,1,2,1,447]
s: 453
i: [1,177,0,7,3,0,0,255] fc24_relat_1
s: 177
i: [1,193,0,7,3,0,0,1] t8_boole
s: 193
d: [2,448,1,2,1,1,40,446,273]
a: [3,458,1,2,1,448]
s: 458
i: [1,171,0,8,2,0,0,1] rd4_relat_1
s: 171
d: [2,461,1,2,1,1,40,460,288]
a: [3,467,1,2,1,461]
s: 467
i: [1,176,0,8,4,0,0,255] cc6_funct_1
s: 176
i: [1,182,0,8,4,0,0,255] cc4_funct_1
s: 182
d: [2,462,1,2,1,1,40,460,273]
a: [3,472,1,2,1,462]
s: 472
d: [2,510,0,6,3,0,21,182,165]
d: [2,511,0,4,2,0,21,510,164]
s: 511
d: [2,513,1,2,1,1,40,511,288]
a: [3,519,1,2,1,513]
s: 519
i: [1,187,0,8,4,0,0,255] cc5_funct_1
s: 187
i: [1,191,0,8,3,0,0,2] t2_subset
s: 191
d: [2,514,1,2,1,1,40,511,273]
a: [3,524,1,2,1,514]
s: 524
i: [1,204,0,8,3,0,0,1] fc17_relat_1
s: 204
d: [2,431,1,5,2,0,40,167,165]
s: 431
d: [2,529,0,6,2,0,21,204,165]
s: 529
i: [1,205,0,8,3,0,0,1] fc17_relat_1
s: 205
d: [2,530,3,2,1,1,40,529,444]
a: [3,541,3,2,1,530]
s: 541
i: [1,207,0,8,3,0,0,255] fc23_relat_1
s: 207
d: [2,432,1,5,2,0,40,167,164]
s: 432
i: [1,208,0,8,3,0,0,255] fc23_relat_1
s: 208
i: [1,209,0,8,3,0,0,1] fc16_relat_1
s: 209
d: [2,545,3,4,2,1,40,209,444]
d: [2,554,3,2,1,2,21,545,263]
a: [3,558,3,2,1,554]
s: 558
i: [1,210,0,8,3,0,0,1] fc16_relat_1
s: 210
d: [2,433,1,5,2,0,40,167,163]
s: 433
i: [1,211,0,8,3,0,0,1] fc8_funct_1
s: 211
i: [1,212,0,8,3,0,0,1] fc8_funct_1
s: 212
d: [2,449,1,5,2,0,40,446,167]
s: 449
i: [1,168,0,9,4,0,0,1] cc8_relat_1
s: 168
d: [2,563,1,6,3,1,40,168,258]
d: [2,564,1,4,2,2,21,563,541]
d: [2,565,1,2,1,3,21,564,263]
a: [3,569,1,2,1,565]
s: 569
i: [1,172,0,9,4,0,0,255] fc18_funct_1
s: 172
i: [1,216,0,9,3,0,0,1] fc26_relat_1
s: 216
d: [2,463,1,5,2,0,40,460,167]
s: 463
i: [1,206,0,10,4,0,0,255] fc19_funct_1
s: 206
d: [2,487,1,5,2,0,40,161,199]
s: 487
i: [1,217,0,10,3,0,0,1] fc26_relat_1
s: 217
i: [1,195,0,11,3,0,0,1] t57_relat_1
s: 195
d: [2,583,1,5,2,1,40,195,248]
d: [2,586,1,3,1,2,21,583,263]
a: [3,590,1,3,1,586]
s: 590
i: [1,201,0,11,2,0,0,1] rd5_relat_1
s: 201
d: [2,592,2,3,1,1,40,590,202]
a: [3,598,2,3,1,592]
s: 598
i: [1,196,0,12,3,0,0,1] t57_relat_1
s: 196
i: [1,213,0,13,5,0,0,255] fc17_funct_1
s: 213
d: [2,593,2,3,1,1,40,590,190]
a: [3,603,2,3,1,593]
s: 603
i: [1,214,0,13,5,0,0,255] fc17_funct_1
s: 214
d: [2,626,3,4,2,1,40,603,162]
a: [3,631,3,2,1,626]
s: 631
i: [1,215,0,13,4,0,0,1] dt_k7_partfun1
s: 215
i: [1,197,0,15,4,0,0,1] t57_relat_1
s: 197
d: [2,615,1,6,2,1,40,196,248]
d: [2,619,1,4,1,2,21,615,263]
a: [3,623,1,4,1,619]
s: 623
i: [1,203,0,16,6,0,0,2] fc11_funct_1
s: 203
d: [2,649,2,3,1,1,40,623,194]
a: [3,663,2,3,1,649]
s: 663
i: [1,192,0,19,5,0,0,1] d6_partfun1
s: 192
i: [1,198,0,19,4,0,0,2] t47_funct_1
s: 198
d: [2,647,2,4,1,1,40,623,202]
a: [3,653,2,4,1,647]
s: 653
d: [2,498,1,5,2,0,40,162,199]
s: 498
d: [2,648,2,4,1,1,40,623,190]
a: [3,658,2,4,1,648]
s: 658
d: [2,515,1,5,2,0,40,511,167]
s: 515
d: [2,526,1,6,2,0,40,191,199]
s: 526
d: [2,709,3,12,5,1,40,658,203]
d: [2,713,3,10,4,2,21,709,253]
d: [2,714,3,8,3,3,21,713,263]
d: [2,715,3,6,2,4,21,714,541]
a: [3,719,3,2,1,715]
s: 719
d: [2,531,1,6,2,0,40,529,511]
s: 531
d: [2,711,3,5,2,1,40,658,162]
a: [3,728,3,3,1,711]
s: 728
d: [2,532,1,6,2,0,40,529,460]
s: 532
d: [2,533,1,6,2,0,40,529,446]
s: 533
d: [2,666,2,22,5,1,40,203,479]
d: [2,668,2,18,4,2,21,666,444]
a: [3,672,2,4,1,668]
s: 672
d: [2,743,3,0,0,2,52,263,672,200]
e: 743
d: [2,744,3,2,1,1,40,672,200]
d: [2,746,3,0,0,2,21,744,263]
e: 746
r: 672
d: [2,534,1,6,2,0,40,529,165]
s: 534
a: [3,671,2,4,1,668]
s: 671
d: [2,535,1,6,2,0,40,529,164]
s: 535
d: [2,536,1,6,2,0,40,529,163]
s: 536
a: [3,680,2,4,1,668]
s: 680
d: [2,754,3,0,0,3,52,263,253,680,212]
e: 754
d: [2,756,3,4,2,1,40,680,212]
d: [2,758,3,2,1,2,21,756,253]
d: [2,759,3,0,0,3,21,758,263]
e: 759
r: 680
d: [2,687,1,29,5,1,44,243,192]
d: [2,690,1,23,4,2,21,687,248]
a: [3,698,1,10,1,690]
s: 698
a: [3,676,2,4,1,668]
s: 676
d: [2,427,1,7,3,0,44,166,166]
s: 427
d: [2,546,1,8,3,0,40,209,511]
s: 546
a: [3,679,2,4,1,668]
s: 679
d: [2,547,1,8,3,0,40,209,460]
s: 547
d: [2,761,2,18,4,1,44,698,198]
d: [2,762,2,12,3,2,21,761,248]
d: [2,763,2,10,2,3,21,762,253]
d: [2,764,2,8,1,4,21,763,263]
a: [3,768,2,8,1,764]
s: 768
d: [2,778,3,0,0,5,52,263,253,258,623,768,192]
e: 778
d: [2,779,3,18,5,1,44,768,192]
d: [2,780,3,11,4,1,19,779]
d: [2,782,3,7,3,2,21,780,623]
d: [2,783,3,5,2,3,21,782,253]
d: [2,784,3,2,1,4,21,783,258]
d: [2,785,3,0,0,5,21,784,263]
e: 785
r: 698
r: 768
a: [3,694,1,5,1,690]
s: 694
d: [2,787,2,0,0,3,52,263,258,694,217]
e: 787
d: [2,788,2,5,2,1,40,694,217]
d: [2,790,2,2,1,2,21,788,258]
d: [2,791,2,0,0,3,21,790,263]
e: 791
f: 791,785,764,759,746,690,619,153,154,155,156,157
% Refutation found. Thanks to Tanya!
% SZS status Theorem for partfun1__t80_partfun1
% SZS output start Proof for partfun1__t80_partfun1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((v1_funct_1(X3) & v5_relat_1(X3,X0) & v1_relat_1(X3)) => (r2_hidden(X1,k9_xtuple_0(k5_relat_1(X3,X2))) => k7_partfun1(X0,k5_relat_1(X3,X2),X1) = k7_partfun1(X0,X3,X1))) [input t80_partfun1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ((v1_funct_1(X3) & v5_relat_1(X3,X0) & v1_relat_1(X3)) => (r2_hidden(X1,k9_xtuple_0(k5_relat_1(X3,X2))) => k7_partfun1(X0,k5_relat_1(X3,X2),X1) = k7_partfun1(X0,X3,X1))) [negated conjecture 1]
16. ! [X0] : ! [X1] : ((v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (r2_hidden(X2,k9_xtuple_0(X1)) => k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2))) [input d6_partfun1]
19. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k5_relat_1(X0,X1))) [input dt_k5_relat_1]
35. ! [X0,X1,X2] : ((v5_relat_1(X2,X1) & v1_relat_1(X2)) => (v5_relat_1(k5_relat_1(X2,X0),X1) & v1_relat_1(k5_relat_1(X2,X0)))) [input fc26_relat_1]
36. ! [X0,X1] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1)))) [input fc8_funct_1]
52. ! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) => k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0))) [input t47_funct_1]
53. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) <=> (r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1)))) [input t57_relat_1]
57. ~! [X0,X1,X2,X3] : ((v1_funct_1(X3) & v5_relat_1(X3,X0) & v1_relat_1(X3)) => (r2_hidden(X1,k9_xtuple_0(k5_relat_1(X3,X2))) => k7_partfun1(X0,k5_relat_1(X3,X2),X1) = k7_partfun1(X0,X3,X1))) [flattening 2]
60. ! [X0,X1] : ((v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (r2_hidden(X2,k9_xtuple_0(X1)) => k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2))) [flattening 16]
63. ! [X0,X1,X2] : (v1_relat_1(X2) => (r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) <=> (r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1)))) [flattening 53]
64. ! [X0,X1,X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) => k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0))) [flattening 52]
67. ? [X0,X1,X2,X3] : ((k7_partfun1(X0,k5_relat_1(X3,X2),X1) != k7_partfun1(X0,X3,X1) & r2_hidden(X1,k9_xtuple_0(k5_relat_1(X3,X2)))) & (v1_funct_1(X3) & v5_relat_1(X3,X0) & v1_relat_1(X3))) [ennf transformation 57]
68. ? [X0,X1,X2,X3] : (k7_partfun1(X0,k5_relat_1(X3,X2),X1) != k7_partfun1(X0,X3,X1) & r2_hidden(X1,k9_xtuple_0(k5_relat_1(X3,X2))) & v1_funct_1(X3) & v5_relat_1(X3,X0) & v1_relat_1(X3)) [flattening 67]
99. ! [X0,X1] : (! [X2] : (k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2) | ~r2_hidden(X2,k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 60]
100. ! [X0,X1] : (! [X2] : (k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2) | ~r2_hidden(X2,k9_xtuple_0(X1))) | ~v1_funct_1(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 99]
103. ! [X0,X1,X2] : ((r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) <=> (r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1))) | ~v1_relat_1(X2)) [ennf transformation 63]
104. ! [X0,X1,X2] : ((k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1)))) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) [ennf transformation 64]
105. ! [X0,X1,X2] : (k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) [flattening 104]
106. ! [X0,X1] : (v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 19]
119. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 36]
120. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 119]
125. ! [X0,X1,X2] : ((v5_relat_1(k5_relat_1(X2,X0),X1) & v1_relat_1(k5_relat_1(X2,X0))) | (~v5_relat_1(X2,X1) | ~v1_relat_1(X2))) [ennf transformation 35]
126. ! [X0,X1,X2] : ((v5_relat_1(k5_relat_1(X2,X0),X1) & v1_relat_1(k5_relat_1(X2,X0))) | ~v5_relat_1(X2,X1) | ~v1_relat_1(X2)) [flattening 125]
127. ? [X0,X1,X2,X3] : (k7_partfun1(X0,k5_relat_1(X3,X2),X1) != k7_partfun1(X0,X3,X1) & r2_hidden(X1,k9_xtuple_0(k5_relat_1(X3,X2))) & v1_funct_1(X3) & v5_relat_1(X3,X0) & v1_relat_1(X3)) => (k7_partfun1(sK0,k5_relat_1(sK3,sK2),sK1) != k7_partfun1(sK0,sK3,sK1) & r2_hidden(sK1,k9_xtuple_0(k5_relat_1(sK3,sK2))) & v1_funct_1(sK3) & v5_relat_1(sK3,sK0) & v1_relat_1(sK3)) [choice axiom]
128. k7_partfun1(sK0,k5_relat_1(sK3,sK2),sK1) != k7_partfun1(sK0,sK3,sK1) & r2_hidden(sK1,k9_xtuple_0(k5_relat_1(sK3,sK2))) & v1_funct_1(sK3) & v5_relat_1(sK3,sK0) & v1_relat_1(sK3) [skolemisation 68,127]
129. ! [X0,X1,X2] : (((r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | (~r2_hidden(X0,k9_xtuple_0(X2)) | ~r2_hidden(X0,X1))) & ((r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1)) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))))) | ~v1_relat_1(X2)) [nnf transformation 103]
130. ! [X0,X1,X2] : (((r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | ~r2_hidden(X0,k9_xtuple_0(X2)) | ~r2_hidden(X0,X1)) & ((r2_hidden(X0,k9_xtuple_0(X2)) & r2_hidden(X0,X1)) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))))) | ~v1_relat_1(X2)) [flattening 129]
153. v1_relat_1(sK3) [cnf transformation 128] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
154. v5_relat_1(sK3,sK0) [cnf transformation 128] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
155. v1_funct_1(sK3) [cnf transformation 128] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
156. r2_hidden(sK1,k9_xtuple_0(k5_relat_1(sK3,sK2))) [cnf transformation 128] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
157. k7_partfun1(sK0,k5_relat_1(sK3,sK2),sK1) != k7_partfun1(sK0,sK3,sK1) [cnf transformation 128] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
192. k7_partfun1(X0,X1,X2) = k1_funct_1(X1,X2) | ~r2_hidden(X2,k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v5_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 100] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
196. ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | r2_hidden(X0,k9_xtuple_0(X2)) | ~v1_relat_1(X2) [cnf transformation 130] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
198. k1_funct_1(k5_relat_1(X2,X1),X0) = k1_funct_1(X2,X0) | ~r2_hidden(X0,k9_xtuple_0(k5_relat_1(X2,X1))) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [cnf transformation 105] {a:0,w:19,wCS:7600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
200. v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 106] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
212. v1_funct_1(k5_relat_1(X0,X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 120] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
217. v5_relat_1(k5_relat_1(X2,X0),X1) | ~v5_relat_1(X2,X1) | ~v1_relat_1(X2) [cnf transformation 126] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
241. 1 <=> k7_partfun1(sK0,k5_relat_1(sK3,sK2),sK1) = k7_partfun1(sK0,sK3,sK1) [avatar definition]
243. k7_partfun1(sK0,k5_relat_1(sK3,sK2),sK1) != k7_partfun1(sK0,sK3,sK1) <- (~1) [avatar component clause 241] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
244. ~1 [avatar split clause 157,241]
246. 2 <=> r2_hidden(sK1,k9_xtuple_0(k5_relat_1(sK3,sK2))) [avatar definition]
248. r2_hidden(sK1,k9_xtuple_0(k5_relat_1(sK3,sK2))) <- (2) [avatar component clause 246] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
249. 2 [avatar split clause 156,246]
251. 3 <=> v1_funct_1(sK3) [avatar definition]
253. v1_funct_1(sK3) <- (3) [avatar component clause 251] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
254. 3 [avatar split clause 155,251]
256. 4 <=> v5_relat_1(sK3,sK0) [avatar definition]
258. v5_relat_1(sK3,sK0) <- (4) [avatar component clause 256] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
259. 4 [avatar split clause 154,256]
261. 5 <=> v1_relat_1(sK3) [avatar definition]
263. v1_relat_1(sK3) <- (5) [avatar component clause 261] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
264. 5 [avatar split clause 153,261]
615. r2_hidden(sK1,k9_xtuple_0(sK3)) | ~v1_relat_1(sK3) <- (2) [resolution 196,248] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
619. r2_hidden(sK1,k9_xtuple_0(sK3)) <- (2, 5) [subsumption resolution 615,263] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
621. 56 <=> r2_hidden(sK1,k9_xtuple_0(sK3)) [avatar definition]
623. r2_hidden(sK1,k9_xtuple_0(sK3)) <- (56) [avatar component clause 621] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
624. 56 | ~2 | ~5 [avatar split clause 619,261,246,621]
670. 61 <=> v1_relat_1(k5_relat_1(sK3,sK2)) [avatar definition]
672. ~v1_relat_1(k5_relat_1(sK3,sK2)) <- (~61) [avatar component clause 670] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
678. 63 <=> v1_funct_1(k5_relat_1(sK3,sK2)) [avatar definition]
680. ~v1_funct_1(k5_relat_1(sK3,sK2)) <- (~63) [avatar component clause 678] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
687. k7_partfun1(sK0,sK3,sK1) != k1_funct_1(k5_relat_1(sK3,sK2),sK1) | ~r2_hidden(sK1,k9_xtuple_0(k5_relat_1(sK3,sK2))) | ~v1_funct_1(k5_relat_1(sK3,sK2)) | ~v5_relat_1(k5_relat_1(sK3,sK2),sK0) | ~v1_relat_1(k5_relat_1(sK3,sK2)) <- (~1) [superposition 243,192] {a:1,w:29,wCS:2900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
690. k7_partfun1(sK0,sK3,sK1) != k1_funct_1(k5_relat_1(sK3,sK2),sK1) | ~v1_funct_1(k5_relat_1(sK3,sK2)) | ~v5_relat_1(k5_relat_1(sK3,sK2),sK0) | ~v1_relat_1(k5_relat_1(sK3,sK2)) <- (~1, 2) [subsumption resolution 687,248] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
692. 65 <=> v5_relat_1(k5_relat_1(sK3,sK2),sK0) [avatar definition]
694. ~v5_relat_1(k5_relat_1(sK3,sK2),sK0) <- (~65) [avatar component clause 692] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
696. 66 <=> k7_partfun1(sK0,sK3,sK1) = k1_funct_1(k5_relat_1(sK3,sK2),sK1) [avatar definition]
698. k7_partfun1(sK0,sK3,sK1) != k1_funct_1(k5_relat_1(sK3,sK2),sK1) <- (~66) [avatar component clause 696] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
699. ~61 | ~65 | ~63 | ~66 | 1 | ~2 [avatar split clause 690,246,241,696,678,692,670]
744. ~v1_relat_1(sK3) <- (~61) [resolution 672,200] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
746. $false <- (5, ~61) [subsumption resolution 744,263] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
747. ~5 | 61 [avatar contradiction clause 746]
756. ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (~63) [resolution 680,212] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
758. ~v1_relat_1(sK3) <- (3, ~63) [subsumption resolution 756,253] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
759. $false <- (3, 5, ~63) [subsumption resolution 758,263] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
760. ~3 | ~5 | 63 [avatar contradiction clause 759]
761. k7_partfun1(sK0,sK3,sK1) != k1_funct_1(sK3,sK1) | ~r2_hidden(sK1,k9_xtuple_0(k5_relat_1(sK3,sK2))) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (~66) [superposition 698,198] {a:2,w:18,wCS:1800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
762. k7_partfun1(sK0,sK3,sK1) != k1_funct_1(sK3,sK1) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) <- (2, ~66) [subsumption resolution 761,248] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
763. k7_partfun1(sK0,sK3,sK1) != k1_funct_1(sK3,sK1) | ~v1_relat_1(sK3) <- (2, 3, ~66) [subsumption resolution 762,253] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
764. k7_partfun1(sK0,sK3,sK1) != k1_funct_1(sK3,sK1) <- (2, 3, 5, ~66) [subsumption resolution 763,263] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
766. 70 <=> k7_partfun1(sK0,sK3,sK1) = k1_funct_1(sK3,sK1) [avatar definition]
768. k7_partfun1(sK0,sK3,sK1) != k1_funct_1(sK3,sK1) <- (~70) [avatar component clause 766] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
769. ~70 | ~2 | ~3 | ~5 | 66 [avatar split clause 764,696,261,251,246,766]
779. k1_funct_1(sK3,sK1) != k1_funct_1(sK3,sK1) | ~r2_hidden(sK1,k9_xtuple_0(sK3)) | ~v1_funct_1(sK3) | ~v5_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~70) [superposition 768,192] {a:3,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
780. ~r2_hidden(sK1,k9_xtuple_0(sK3)) | ~v1_funct_1(sK3) | ~v5_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~70) [trivial inequality removal 779] {a:3,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
782. ~v1_funct_1(sK3) | ~v5_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (56, ~70) [subsumption resolution 780,623] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
783. ~v5_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (3, 56, ~70) [subsumption resolution 782,253] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
784. ~v1_relat_1(sK3) <- (3, 4, 56, ~70) [subsumption resolution 783,258] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
785. $false <- (3, 4, 5, 56, ~70) [subsumption resolution 784,263] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
786. ~3 | ~4 | ~5 | ~56 | 70 [avatar contradiction clause 785]
788. ~v5_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~65) [resolution 694,217] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
790. ~v1_relat_1(sK3) <- (4, ~65) [subsumption resolution 788,258] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
791. $false <- (4, 5, ~65) [subsumption resolution 790,263] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
792. ~4 | ~5 | 65 [avatar contradiction clause 791]
793. $false [avatar sat refutation 244,249,254,259,264,624,699,747,760,769,786,792] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:368,thDist:-368}
% SZS output end Proof for partfun1__t80_partfun1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 55
% >>> Preprocessing
% Introduced skolems: 15
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 87
% Generated clauses: 356
% Active clauses: 124
% Passive clauses: 165
% Final active clauses: 119
% Final passive clauses: 30
% Main loop iterations started: 135
% 
% >>> Simplifying Inferences
% Duplicate literals: 18
% Trivial inequalities: 1
% Fw subsumption resolutions: 34
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 34
% Forward subsumptions: 44
% 
% >>> Generating Inferences
% Binary resolution: 93
% Unit resulting resolution: 4
% Forward superposition: 34
% Backward superposition: 15
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 22
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 76
% SAT solver unit clauses: 30
% SAT solver binary clauses: 24
% 
% Memory used [KB]: 9850
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% forward subsumption: 0.001 s
% other: 0.005 s ( own 0.003 s ) 

