i: [1,25,0,2,1,1,0,0] t17_orders_1
a: [3,47,0,2,1,25]
s: 47
i: [1,24,0,2,1,1,0,0] t17_orders_1
a: [3,52,0,2,1,24]
s: 52
i: [1,26,0,3,1,1,0,0] t17_orders_1
a: [3,42,0,3,1,26]
s: 42
i: [1,27,0,5,2,0,0,1] dt_k2_relat_1
s: 27
i: [1,29,0,6,3,0,0,1] d4_orders_1
s: 29
i: [1,30,0,6,3,0,0,1] d4_orders_1
s: 30
d: [2,54,1,4,2,1,40,29,47]
d: [2,55,1,2,1,2,21,54,52]
a: [3,59,1,2,1,55]
s: 59
i: [1,31,0,6,3,0,0,1] d4_orders_1
s: 31
d: [2,61,1,4,2,1,40,30,47]
d: [2,62,1,2,1,2,21,61,52]
a: [3,66,1,2,1,62]
s: 66
i: [1,28,0,7,2,0,0,255] involutiveness_k2_relat_1
s: 28
i: [1,33,0,7,3,0,0,2] fc17_relat_2
s: 33
d: [2,68,1,4,2,1,40,31,47]
d: [2,69,1,2,1,2,21,68,52]
a: [3,73,1,2,1,69]
s: 73
i: [1,34,0,7,3,0,0,2] fc17_relat_2
s: 34
i: [1,35,0,7,3,0,0,2] fc2_relat_2
s: 35
i: [1,36,0,7,3,0,0,2] fc2_relat_2
s: 36
i: [1,37,0,7,3,0,0,2] fc21_relat_2
s: 37
i: [1,38,0,7,3,0,0,2] fc21_relat_2
s: 38
i: [1,32,0,10,5,0,0,1] d4_orders_1
s: 32
d: [2,76,1,10,4,0,44,34,28]
d: [2,77,1,7,3,0,21,76,27]
s: 77
d: [2,78,1,10,4,0,44,36,28]
d: [2,79,1,7,3,0,21,78,27]
s: 79
d: [2,80,1,10,4,0,44,38,28]
d: [2,81,1,7,3,0,21,80,27]
s: 81
d: [2,83,1,16,6,0,40,32,34]
d: [2,84,1,13,5,0,21,83,27]
s: 84
d: [2,95,2,14,6,0,44,84,28]
d: [2,97,2,11,5,0,21,95,27]
s: 97
d: [2,94,2,14,6,0,40,84,38]
d: [2,96,2,12,5,0,18,94]
s: 96
d: [2,101,3,13,6,0,40,96,36]
d: [2,103,3,11,5,0,18,101]
s: 103
d: [2,105,4,0,0,5,52,59,66,52,73,42,103]
e: 105
d: [2,106,4,8,4,1,40,103,42]
d: [2,112,4,6,3,2,21,106,73]
d: [2,113,4,4,2,3,21,112,52]
d: [2,114,4,2,1,4,21,113,66]
d: [2,115,4,0,0,5,21,114,59]
e: 115
f: 115,69,62,55,24,25,26
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t17_orders_1
% SZS output start Proof for orders_1__t17_orders_1
1. ! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) => v2_orders_1(k2_relat_1(X0)))) [input t17_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) => v2_orders_1(k2_relat_1(X0)))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d4_orders_1]
4. ! [X0] : (v1_relat_1(X0) => v1_relat_1(k2_relat_1(X0))) [input dt_k2_relat_1]
5. ! [X0] : ((v4_relat_2(X0) & v1_relat_1(X0)) => (v4_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0)))) [input fc17_relat_2]
6. ! [X0] : ((v8_relat_2(X0) & v1_relat_1(X0)) => (v8_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0)))) [input fc21_relat_2]
7. ! [X0] : ((v1_relat_2(X0) & v1_relat_1(X0)) => (v1_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0)))) [input fc2_relat_2]
9. ? [X0] : ((~v2_orders_1(k2_relat_1(X0)) & v2_orders_1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
10. ? [X0] : (~v2_orders_1(k2_relat_1(X0)) & v2_orders_1(X0) & v1_relat_1(X0)) [flattening 9]
11. ! [X0] : (v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 4]
13. ! [X0] : ((v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
14. ! [X0] : ((v4_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | (~v4_relat_2(X0) | ~v1_relat_1(X0))) [ennf transformation 5]
15. ! [X0] : ((v4_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | ~v4_relat_2(X0) | ~v1_relat_1(X0)) [flattening 14]
16. ! [X0] : ((v1_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | (~v1_relat_2(X0) | ~v1_relat_1(X0))) [ennf transformation 7]
17. ! [X0] : ((v1_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | ~v1_relat_2(X0) | ~v1_relat_1(X0)) [flattening 16]
18. ! [X0] : ((v8_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | (~v8_relat_2(X0) | ~v1_relat_1(X0))) [ennf transformation 6]
19. ! [X0] : ((v8_relat_2(k2_relat_1(X0)) & v1_relat_1(k2_relat_1(X0))) | ~v8_relat_2(X0) | ~v1_relat_1(X0)) [flattening 18]
20. ? [X0] : (~v2_orders_1(k2_relat_1(X0)) & v2_orders_1(X0) & v1_relat_1(X0)) => (~v2_orders_1(k2_relat_1(sK0)) & v2_orders_1(sK0) & v1_relat_1(sK0)) [choice axiom]
21. ~v2_orders_1(k2_relat_1(sK0)) & v2_orders_1(sK0) & v1_relat_1(sK0) [skolemisation 10,20]
22. ! [X0] : (((v2_orders_1(X0) | (~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 13]
23. ! [X0] : (((v2_orders_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 22]
24. v1_relat_1(sK0) [cnf transformation 21] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
25. v2_orders_1(sK0) [cnf transformation 21] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
26. ~v2_orders_1(k2_relat_1(sK0)) [cnf transformation 21] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 11] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
29. ~v2_orders_1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
30. ~v2_orders_1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
31. ~v2_orders_1(X0) | v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
32. ~v4_relat_2(X0) | v2_orders_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. v4_relat_2(k2_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 15] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
36. v1_relat_2(k2_relat_1(X0)) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 17] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
38. v8_relat_2(k2_relat_1(X0)) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 19] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
40. 1 <=> v2_orders_1(k2_relat_1(sK0)) [avatar definition]
42. ~v2_orders_1(k2_relat_1(sK0)) <- (~1) [avatar component clause 40] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
43. ~1 [avatar split clause 26,40]
45. 2 <=> v2_orders_1(sK0) [avatar definition]
47. v2_orders_1(sK0) <- (2) [avatar component clause 45] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
48. 2 [avatar split clause 25,45]
50. 3 <=> v1_relat_1(sK0) [avatar definition]
52. v1_relat_1(sK0) <- (3) [avatar component clause 50] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
53. 3 [avatar split clause 24,50]
54. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 29,47] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
55. v1_relat_2(sK0) <- (2, 3) [subsumption resolution 54,52] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
57. 4 <=> v1_relat_2(sK0) [avatar definition]
59. v1_relat_2(sK0) <- (4) [avatar component clause 57] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
60. 4 | ~2 | ~3 [avatar split clause 55,50,45,57]
61. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 30,47] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
62. v8_relat_2(sK0) <- (2, 3) [subsumption resolution 61,52] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
64. 5 <=> v8_relat_2(sK0) [avatar definition]
66. v8_relat_2(sK0) <- (5) [avatar component clause 64] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
67. 5 | ~2 | ~3 [avatar split clause 62,50,45,64]
68. v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 31,47] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
69. v4_relat_2(sK0) <- (2, 3) [subsumption resolution 68,52] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
71. 6 <=> v4_relat_2(sK0) [avatar definition]
73. v4_relat_2(sK0) <- (6) [avatar component clause 71] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
74. 6 | ~2 | ~3 [avatar split clause 69,50,45,71]
83. v2_orders_1(k2_relat_1(X0)) | ~v8_relat_2(k2_relat_1(X0)) | ~v1_relat_2(k2_relat_1(X0)) | ~v1_relat_1(k2_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [resolution 32,34] {a:1,w:16,wCS:6400,sine:1,thAx:0,allAx:2,thDist:-2}
84. ~v8_relat_2(k2_relat_1(X0)) | v2_orders_1(k2_relat_1(X0)) | ~v1_relat_2(k2_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [subsumption resolution 83,27] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
94. v2_orders_1(k2_relat_1(X0)) | ~v1_relat_2(k2_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [resolution 84,38] {a:2,w:14,wCS:5600,sine:1,thAx:0,allAx:3,thDist:-3}
96. ~v1_relat_2(k2_relat_1(X0)) | v2_orders_1(k2_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) | ~v8_relat_2(X0) [duplicate literal removal 94] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
101. v2_orders_1(k2_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [resolution 96,36] {a:3,w:13,wCS:5200,sine:1,thAx:0,allAx:4,thDist:-4}
103. v2_orders_1(k2_relat_1(X0)) | ~v4_relat_2(X0) | ~v1_relat_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) [duplicate literal removal 101] {a:3,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
106. ~v4_relat_2(sK0) | ~v1_relat_1(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1) [resolution 103,42] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
112. ~v1_relat_1(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1, 6) [subsumption resolution 106,73] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
113. ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1, 3, 6) [subsumption resolution 112,52] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
114. ~v1_relat_2(sK0) <- (~1, 3, 5, 6) [subsumption resolution 113,66] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
115. $false <- (~1, 3, 4, 5, 6) [subsumption resolution 114,59] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
116. 1 | ~3 | ~4 | ~5 | ~6 [avatar contradiction clause 115]
118. $false [avatar sat refutation 43,48,53,60,67,74,116] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:50,thDist:-50}
% SZS output end Proof for orders_1__t17_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 7
% Introduced skolems: 1
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 68
% Active clauses: 22
% Passive clauses: 26
% Final active clauses: 22
% Final passive clauses: 1
% Main loop iterations started: 25
% 
% >>> Simplifying Inferences
% Duplicate literals: 6
% Fw subsumption resolutions: 14
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 15
% Unit resulting resolution: 1
% Forward superposition: 10
% Backward superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

