i: [1,51,0,2,1,1,0,0] t8_funct_4
a: [3,80,0,2,1,51]
s: 80
i: [1,50,0,2,1,1,0,0] t8_funct_4
a: [3,85,0,2,1,50]
s: 85
i: [1,52,0,7,1,1,0,0] t8_funct_4
a: [3,75,0,7,1,52]
s: 75
i: [1,71,0,2,1,0,0,1] rc1_funct_1
a: [3,90,0,2,1,71]
s: 90
i: [1,70,0,2,1,0,0,1] rc1_funct_1
a: [3,95,0,2,1,70]
s: 95
i: [1,61,0,3,1,0,0,1] reflexivity_r1_tarski
s: 61
i: [1,60,0,4,1,0,0,255] existence_m1_subset_1
s: 60
i: [1,62,0,6,2,0,0,1] dt_k6_relat_1
s: 62
i: [1,63,0,6,2,0,0,1] dt_k5_relat_1
s: 63
i: [1,55,0,7,2,0,0,1] t86_relat_1
s: 55
i: [1,56,0,7,2,0,0,1] t59_relat_1
s: 56
i: [1,57,0,7,2,0,0,255] t3_subset
s: 57
i: [1,58,0,7,2,0,0,255] t3_subset
s: 58
d: [2,97,1,5,1,0,40,57,60]
s: 97
i: [1,53,0,8,3,0,0,255] cc2_relat_1
s: 53
i: [1,66,0,8,3,0,0,1] fc8_funct_1
s: 66
d: [2,99,1,6,2,0,40,53,60]
s: 99
i: [1,67,0,8,3,0,0,1] fc8_funct_1
s: 67
d: [2,100,1,7,3,0,40,53,58]
s: 100
i: [1,68,0,8,3,0,0,1] fc9_funct_1
s: 68
i: [1,69,0,8,3,0,0,1] fc9_funct_1
s: 69
i: [1,59,0,9,3,0,0,1] t1_xboole_1
s: 59
i: [1,54,0,10,4,0,0,255] cc3_funct_1
s: 54
d: [2,111,1,8,3,0,40,54,60]
s: 111
i: [1,64,0,11,2,0,0,255] rd7_relat_1
s: 64
i: [1,65,0,11,2,0,0,255] rd5_relat_1
s: 65
d: [2,110,2,8,2,0,40,59,97]
s: 110
d: [2,112,1,9,4,0,40,54,58]
s: 112
d: [2,124,3,7,1,0,40,110,97]
s: 124
d: [2,108,1,10,3,0,40,59,55]
s: 108
d: [2,109,1,10,3,0,40,59,56]
s: 109
d: [2,133,4,8,2,0,40,124,100]
s: 133
d: [2,138,2,9,2,0,40,108,97]
s: 138
d: [2,153,3,10,3,0,40,138,100]
d: [2,158,3,8,2,0,18,153]
s: 158
d: [2,136,2,13,3,0,40,108,55]
d: [2,141,2,9,2,0,21,136,62]
s: 141
d: [2,137,2,13,3,0,40,108,56]
d: [2,142,2,9,2,0,21,137,62]
s: 142
d: [2,176,3,0,0,2,52,85,75,142]
e: 176
d: [2,177,3,2,1,1,40,142,75]
d: [2,188,3,0,0,2,21,177,85]
e: 188
f: 188,50,52
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t8_funct_4
% SZS output start Proof for funct_4__t8_funct_4
1. ! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),X2)) [input t8_funct_4]
2. ~! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),X2)) [negated conjecture 1]
7. ! [X0,X1] : (v1_relat_1(X1) => v1_relat_1(k6_relat_1(X0,X1))) [input dt_k6_relat_1]
16. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
18. ! [X0] : ! [X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [input t59_relat_1]
19. ! [X0] : ! [X1] : (v1_relat_1(X1) => r1_tarski(k6_relat_1(X0,X1),X1)) [input t86_relat_1]
20. ~! [X0,X1,X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),X2)) [flattening 2]
21. ! [X0,X1] : (v1_relat_1(X1) => r1_tarski(k6_relat_1(X0,X1),X1)) [flattening 19]
22. ! [X0,X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [flattening 18]
24. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 16]
26. ? [X0,X1,X2] : (~r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),X2) & (v1_funct_1(X2) & v1_relat_1(X2))) [ennf transformation 20]
27. ? [X0,X1,X2] : (~r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),X2) & v1_funct_1(X2) & v1_relat_1(X2)) [flattening 26]
31. ! [X0,X1] : (r1_tarski(k6_relat_1(X0,X1),X1) | ~v1_relat_1(X1)) [ennf transformation 21]
32. ! [X0,X1] : (r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1)) [ennf transformation 22]
33. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 24]
34. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 33]
35. ! [X0,X1] : (v1_relat_1(k6_relat_1(X0,X1)) | ~v1_relat_1(X1)) [ennf transformation 7]
43. ? [X0,X1,X2] : (~r1_tarski(k5_relat_1(k6_relat_1(X0,X2),X1),X2) & v1_funct_1(X2) & v1_relat_1(X2)) => (~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),sK2) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
44. ~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),sK2) & v1_funct_1(sK2) & v1_relat_1(sK2) [skolemisation 27,43]
50. v1_relat_1(sK2) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. ~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),sK2) [cnf transformation 44] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. r1_tarski(k6_relat_1(X0,X1),X1) | ~v1_relat_1(X1) [cnf transformation 31] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1) [cnf transformation 32] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 34] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. v1_relat_1(k6_relat_1(X0,X1)) | ~v1_relat_1(X1) [cnf transformation 35] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. 1 <=> r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),sK2) [avatar definition]
75. ~r1_tarski(k5_relat_1(k6_relat_1(sK0,sK2),sK1),sK2) <- (~1) [avatar component clause 73] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. ~1 [avatar split clause 52,73]
83. 3 <=> v1_relat_1(sK2) [avatar definition]
85. v1_relat_1(sK2) <- (3) [avatar component clause 83] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
86. 3 [avatar split clause 50,83]
108. ~r1_tarski(X2,k6_relat_1(X4,X3)) | r1_tarski(X2,X3) | ~v1_relat_1(X3) [resolution 59,55] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
137. r1_tarski(k5_relat_1(k6_relat_1(X5,X6),X7),X6) | ~v1_relat_1(X6) | ~v1_relat_1(k6_relat_1(X5,X6)) [resolution 108,56] {a:2,w:13,wCS:5200,sine:1,thAx:0,allAx:3,thDist:-3}
142. r1_tarski(k5_relat_1(k6_relat_1(X5,X6),X7),X6) | ~v1_relat_1(X6) [subsumption resolution 137,62] {a:2,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
177. ~v1_relat_1(sK2) <- (~1) [resolution 142,75] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
188. $false <- (~1, 3) [subsumption resolution 177,85] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
189. 1 | ~3 [avatar contradiction clause 188]
192. $false [avatar sat refutation 76,86,189] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:27,thDist:-27}
% SZS output end Proof for funct_4__t8_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 22
% Generated clauses: 120
% Active clauses: 34
% Passive clauses: 57
% Final active clauses: 34
% Final passive clauses: 21
% Main loop iterations started: 36
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Forward subsumptions: 20
% 
% >>> Generating Inferences
% Binary resolution: 55
% Unit resulting resolution: 1
% Forward superposition: 7
% Backward superposition: 6
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 5
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

