i: [1,50,0,8,1,1,0,0] t16_zfmisc_1
a: [3,84,0,8,1,50]
s: 84
i: [1,51,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,89,0,2,1,51]
s: 89
i: [1,52,0,2,1,0,0,1] fc1_xboole_0
a: [3,94,0,2,1,52]
s: 94
i: [1,74,0,2,1,0,0,2] rc2_xboole_0
a: [3,104,0,2,1,74]
s: 104
i: [1,75,0,2,1,0,0,2] rc1_xboole_0
a: [3,109,0,2,1,75]
s: 109
i: [1,54,0,3,1,0,0,1] fc2_xboole_0
s: 54
i: [1,70,0,3,1,0,0,2] reflexivity_r1_tarski
s: 70
i: [1,53,0,3,1,0,0,255] d2_xboole_0
a: [3,99,0,3,1,53]
s: 99
r: 89
i: [1,71,0,4,1,0,0,1] fc3_xboole_0
s: 71
i: [1,55,0,5,1,0,0,1] t4_boole
s: 55
i: [1,56,0,5,1,0,0,1] t3_boole
s: 56
i: [1,57,0,5,2,0,0,1] t6_boole
s: 57
i: [1,63,0,5,2,0,0,2] t7_boole
s: 63
i: [1,77,0,5,1,0,0,1] d2_tarski
s: 77
i: [1,79,0,5,1,0,0,1] d2_tarski
s: 79
i: [1,73,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 73
d: [2,118,1,5,1,0,40,73,77]
s: 118
i: [1,58,0,7,2,0,0,1] l1_zfmisc_1
s: 58
d: [2,120,1,4,1,0,40,58,70]
s: 120
i: [1,59,0,7,2,0,0,1] l1_zfmisc_1
s: 59
i: [1,62,0,7,3,0,0,1] t8_boole
s: 62
d: [2,121,2,4,1,0,40,120,73]
s: 121
i: [1,72,0,7,1,0,0,1] commutativity_k2_tarski
s: 72
d: [2,119,1,5,1,0,40,73,79]
s: 119
i: [1,60,0,8,2,0,0,2] t37_xboole_1
s: 60
i: [1,61,0,8,2,0,0,2] t37_xboole_1
s: 61
d: [2,138,1,9,2,1,44,84,61]
d: [2,141,1,6,1,1,19,138]
a: [3,145,1,6,1,141]
s: 145
d: [2,147,2,0,0,1,52,79,145,59]
e: 147
d: [2,148,2,5,1,1,40,145,59]
d: [2,150,2,0,0,1,21,148,79]
e: 150
f: 150,141,50
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t16_zfmisc_1
% SZS output start Proof for zfmisc_1__t16_zfmisc_1
1. ! [X0] : ! [X1] : k4_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) = k1_xboole_0 [input t16_zfmisc_1]
2. ~! [X0] : ! [X1] : k4_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) = k1_xboole_0 [negated conjecture 1]
5. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
15. ! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [input l1_zfmisc_1]
19. ! [X0] : ! [X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [input t37_xboole_1]
25. ~! [X0,X1] : k4_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) = k1_xboole_0 [flattening 2]
26. ! [X0,X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [flattening 15]
27. ! [X0,X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [flattening 19]
30. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 5]
32. ? [X0,X1] : k4_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) != k1_xboole_0 [ennf transformation 25]
37. ? [X0,X1] : k4_xboole_0(k1_tarski(X0),k2_tarski(X0,X1)) != k1_xboole_0 => k1_xboole_0 != k4_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) [choice axiom]
38. k1_xboole_0 != k4_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) [skolemisation 32,37]
39. ! [X0,X1] : ((r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1)) & (r2_hidden(X0,X1) | ~r1_tarski(k1_tarski(X0),X1))) [nnf transformation 26]
40. ! [X0,X1] : ((k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | k1_xboole_0 != k4_xboole_0(X0,X1))) [nnf transformation 27]
41. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 30]
42. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 41]
43. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 42]
44. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK2(X0,X1,X2) != X1 & sK2(X0,X1,X2) != X0) | ~r2_hidden(sK2(X0,X1,X2),X2)) & (sK2(X0,X1,X2) = X1 | sK2(X0,X1,X2) = X0 | r2_hidden(sK2(X0,X1,X2),X2)))) [choice axiom]
45. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK2(X0,X1,X2) != X1 & sK2(X0,X1,X2) != X0) | ~r2_hidden(sK2(X0,X1,X2),X2)) & (sK2(X0,X1,X2) = X1 | sK2(X0,X1,X2) = X0 | r2_hidden(sK2(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 43,44]
50. k1_xboole_0 != k4_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) [cnf transformation 38] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. r1_tarski(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 39] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1) [cnf transformation 40] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
65. r2_hidden(X4,X2) | X0 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 45] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
78. r2_hidden(X4,X2) | k2_tarski(X4,X1) != X2 [equality resolution 65] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
79. r2_hidden(X4,k2_tarski(X4,X1)) [equality resolution 78] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
82. 1 <=> k1_xboole_0 = k4_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) [avatar definition]
84. k1_xboole_0 != k4_xboole_0(k1_tarski(sK0),k2_tarski(sK0,sK1)) <- (~1) [avatar component clause 82] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
85. ~1 [avatar split clause 50,82]
138. k1_xboole_0 != k1_xboole_0 | ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) <- (~1) [superposition 84,61] {a:1,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
141. ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) <- (~1) [trivial inequality removal 138] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
143. 7 <=> r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) [avatar definition]
145. ~r1_tarski(k1_tarski(sK0),k2_tarski(sK0,sK1)) <- (~7) [avatar component clause 143] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
146. ~7 | 1 [avatar split clause 141,82,143]
148. ~r2_hidden(sK0,k2_tarski(sK0,sK1)) <- (~7) [resolution 145,59] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
150. $false <- (~7) [subsumption resolution 148,79] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
151. 7 [avatar contradiction clause 150]
152. $false [avatar sat refutation 85,146,151] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:30,thDist:-30}
% SZS output end Proof for zfmisc_1__t16_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 67
% Active clauses: 27
% Passive clauses: 40
% Final active clauses: 26
% Final passive clauses: 13
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Trivial inequalities: 2
% Fw subsumption resolutions: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 1
% Forward superposition: 5
% Backward superposition: 14
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 8
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 9338
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

