i: [1,68,0,2,1,1,0,0] t49_setfam_1
a: [3,109,0,2,1,68]
s: 109
i: [1,69,0,4,1,1,0,0] t49_setfam_1
a: [3,104,0,4,1,69]
s: 104
i: [1,70,0,6,2,1,0,0] t49_setfam_1
s: 70
i: [1,71,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,114,0,2,1,71]
s: 114
i: [1,72,0,2,1,0,0,255] fc1_xboole_0
a: [3,119,0,2,1,72]
s: 119
i: [1,93,0,2,1,0,0,3] rc2_setfam_1
a: [3,129,0,2,1,93]
s: 129
i: [1,94,0,2,1,0,0,1] rc2_xboole_0
a: [3,134,0,2,1,94]
s: 134
i: [1,95,0,2,1,0,0,1] rc1_xboole_0
a: [3,139,0,2,1,95]
s: 139
i: [1,97,0,2,1,0,0,2] rc1_setfam_1
a: [3,144,0,2,1,97]
s: 144
i: [1,96,0,2,1,0,0,2] rc1_setfam_1
a: [3,149,0,2,1,96]
s: 149
i: [1,74,0,3,1,0,0,1] fc2_xboole_0
s: 74
i: [1,73,0,3,1,0,0,255] d2_xboole_0
a: [3,124,0,3,1,73]
s: 124
r: 114
i: [1,75,0,4,2,0,0,3] cc2_setfam_1
s: 75
d: [2,152,1,2,1,1,40,75,139]
a: [3,157,1,2,1,152]
s: 157
i: [1,91,0,4,1,0,0,1] existence_m1_subset_1
s: 91
i: [1,99,0,4,1,0,0,1] d1_tarski
s: 99
d: [2,164,1,4,1,0,40,91,70]
a: [3,168,1,4,1,164]
s: 168
i: [1,77,0,5,2,0,0,3] rc3_setfam_1
s: 77
d: [2,170,2,3,1,1,44,91,168]
a: [3,174,2,3,1,170]
s: 174
i: [1,78,0,5,2,0,0,1] fc2_setfam_1
s: 78
i: [1,80,0,5,2,0,0,255] t6_boole
s: 80
d: [2,153,1,2,1,1,40,75,119]
a: [3,162,1,2,1,153]
s: 162
i: [1,90,0,5,2,0,0,2] t7_boole
s: 90
i: [1,76,0,6,2,0,0,3] rc3_setfam_1
s: 76
i: [1,79,0,6,2,0,0,255] dt_o_1_2_setfam_1
s: 79
i: [1,81,0,6,3,0,0,2] cc3_setfam_1
s: 81
d: [2,181,1,6,2,0,40,76,70]
a: [3,189,1,2,1,181]
s: 189
i: [1,83,0,6,2,0,0,2] t1_subset
s: 83
d: [2,200,2,4,2,1,40,189,81]
d: [2,201,2,2,1,2,21,200,109]
a: [3,205,2,2,1,201]
s: 205
i: [1,92,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 92
i: [1,89,0,7,3,0,0,1] t8_boole
s: 89
d: [2,191,1,6,2,0,40,79,70]
d: [2,192,1,4,1,1,21,191,109]
a: [3,196,1,4,1,192]
s: 196
i: [1,100,0,7,2,0,0,1] d1_tarski
s: 100
d: [2,207,1,4,1,0,40,83,99]
s: 207
i: [1,84,0,8,3,0,0,2] t2_subset
s: 84
i: [1,82,0,9,4,0,0,2] cc1_setfam_1
s: 82
d: [2,213,3,5,2,1,40,84,174]
d: [2,217,3,3,1,2,21,213,109]
a: [3,221,3,3,1,217]
s: 221
i: [1,87,0,14,3,0,0,1] d1_tarski
s: 87
d: [2,231,4,3,1,1,40,221,92]
a: [3,237,4,3,1,231]
s: 237
i: [1,88,0,14,3,0,0,1] d1_tarski
s: 88
d: [2,208,1,4,1,0,40,92,99]
s: 208
d: [2,224,2,8,3,0,40,82,207]
d: [2,230,2,5,2,0,21,224,74]
s: 230
d: [2,215,1,6,2,0,40,84,91]
s: 215
d: [2,247,2,8,2,0,40,215,100]
d: [2,249,2,5,1,0,21,247,74]
s: 249
d: [2,212,1,8,3,0,40,84,79]
d: [2,216,1,6,2,0,18,212]
s: 216
d: [2,214,1,8,3,0,40,84,76]
d: [2,223,1,6,2,0,21,214,75]
s: 223
d: [2,255,2,8,2,0,40,216,100]
d: [2,257,2,5,1,0,21,255,74]
s: 257
d: [2,177,1,7,3,0,44,80,80]
s: 177
d: [2,244,2,6,2,0,40,215,92]
s: 244
d: [2,228,1,7,3,0,40,82,91]
s: 228
d: [2,225,1,9,4,0,40,82,79]
d: [2,229,1,7,3,0,18,225]
s: 229
d: [2,252,2,6,2,0,40,216,92]
s: 252
d: [2,241,1,11,3,0,40,87,90]
s: 241
d: [2,258,2,6,2,0,40,223,92]
s: 258
d: [2,239,1,14,3,0,40,87,92]
s: 239
d: [2,240,1,14,3,0,40,87,83]
s: 240
d: [2,277,2,14,3,0,40,240,70]
s: 277
d: [2,242,1,17,3,0,40,87,100]
s: 242
d: [2,286,3,12,3,0,46,277]
s: 286
d: [2,261,2,8,2,0,40,223,100]
s: 261
d: [2,295,2,14,3,0,46,242]
s: 295
d: [2,302,4,9,2,0,47,286]
d: [2,303,4,5,1,1,21,302,104]
a: [3,307,4,5,1,303]
s: 307
d: [2,320,5,10,3,1,44,88,307]
d: [2,322,5,7,2,1,19,320]
d: [2,323,5,3,1,2,21,322,104]
d: [2,324,5,0,0,3,21,323,221]
e: 324
f: 324,303,217,170,164,68,69
% Refutation found. Thanks to Tanya!
% SZS status Theorem for setfam_1__t49_setfam_1
% SZS output start Proof for setfam_1__t49_setfam_1
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ~(! [X2] : (m1_subset_1(X2,X0) => X1 = X2) & k1_tarski(X1) != X0)) [input t49_setfam_1]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : ~(! [X2] : (m1_subset_1(X2,X0) => X1 = X2) & k1_tarski(X1) != X0)) [negated conjecture 1]
7. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
14. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
23. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
24. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
28. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 23]
29. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 24]
30. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 7]
33. ? [X0] : (? [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,X0)) & k1_tarski(X1) != X0) & ~v1_xboole_0(X0)) [ennf transformation 2]
43. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 28]
44. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 29]
45. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 44]
49. ? [X0] : (? [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,X0)) & k1_tarski(X1) != X0) & ~v1_xboole_0(X0)) => (? [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,sK0)) & k1_tarski(X1) != sK0) & ~v1_xboole_0(sK0)) [choice axiom]
50. ? [X1] : (! [X2] : (X1 = X2 | ~m1_subset_1(X2,sK0)) & k1_tarski(X1) != sK0) => (! [X2] : (sK1 = X2 | ~m1_subset_1(X2,sK0)) & sK0 != k1_tarski(sK1)) [choice axiom]
51. (! [X2] : (sK1 = X2 | ~m1_subset_1(X2,sK0)) & sK0 != k1_tarski(sK1)) & ~v1_xboole_0(sK0) [skolemisation 33,50,49]
54. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 30]
55. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 54]
56. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
57. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 55,56]
58. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK4(X0),X0)) [choice axiom]
59. ! [X0] : m1_subset_1(sK4(X0),X0) [skolemisation 14,58]
68. ~v1_xboole_0(sK0) [cnf transformation 51] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
69. sK0 != k1_tarski(sK1) [cnf transformation 51] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
70. ~m1_subset_1(X2,sK0) | sK1 = X2 [cnf transformation 51] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
83. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 43] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
84. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 45] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
87. r2_hidden(sK3(X0,X1),X1) | sK3(X0,X1) = X0 | k1_tarski(X0) = X1 [cnf transformation 57] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
88. sK3(X0,X1) != X0 | k1_tarski(X0) = X1 | ~r2_hidden(sK3(X0,X1),X1) [cnf transformation 57] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
91. m1_subset_1(sK4(X0),X0) [cnf transformation 59] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
102. 1 <=> sK0 = k1_tarski(sK1) [avatar definition]
104. sK0 != k1_tarski(sK1) <- (~1) [avatar component clause 102] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. ~1 [avatar split clause 69,102]
107. 2 <=> v1_xboole_0(sK0) [avatar definition]
109. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 107] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. ~2 [avatar split clause 68,107]
164. sK1 = sK4(sK0) [resolution 91,70] {a:1,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
166. 13 <=> sK1 = sK4(sK0) [avatar definition]
168. sK1 = sK4(sK0) <- (13) [avatar component clause 166] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
169. 13 [avatar split clause 164,166]
170. m1_subset_1(sK1,sK0) <- (13) [superposition 91,168] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
172. 14 <=> m1_subset_1(sK1,sK0) [avatar definition]
174. m1_subset_1(sK1,sK0) <- (14) [avatar component clause 172] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
175. 14 | ~13 [avatar split clause 170,166,172]
213. v1_xboole_0(sK0) | r2_hidden(sK1,sK0) <- (14) [resolution 84,174] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
217. r2_hidden(sK1,sK0) <- (~2, 14) [subsumption resolution 213,109] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
219. 19 <=> r2_hidden(sK1,sK0) [avatar definition]
221. r2_hidden(sK1,sK0) <- (19) [avatar component clause 219] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
222. 19 | 2 | ~14 [avatar split clause 217,172,107,219]
240. m1_subset_1(sK3(X2,X3),X3) | k1_tarski(X2) = X3 | sK3(X2,X3) = X2 [resolution 87,83] {a:1,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
277. sK3(X0,sK0) = X0 | k1_tarski(X0) = sK0 | sK1 = sK3(X0,sK0) [resolution 240,70] {a:2,w:14,wCS:1400,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
286. sK1 != X0 | k1_tarski(X0) = sK0 | sK1 = sK3(X0,sK0) [equality factoring 277] {a:3,w:12,wCS:1200,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
302. sK0 = k1_tarski(sK1) | sK1 = sK3(sK1,sK0) [equality resolution 286] {a:4,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
303. sK1 = sK3(sK1,sK0) <- (~1) [subsumption resolution 302,104] {a:4,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
305. 21 <=> sK1 = sK3(sK1,sK0) [avatar definition]
307. sK1 = sK3(sK1,sK0) <- (21) [avatar component clause 305] {a:4,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
308. 21 | 1 [avatar split clause 303,102,305]
320. sK1 != sK1 | sK0 = k1_tarski(sK1) | ~r2_hidden(sK1,sK0) <- (21) [superposition 88,307] {a:5,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
322. sK0 = k1_tarski(sK1) | ~r2_hidden(sK1,sK0) <- (21) [trivial inequality removal 320] {a:5,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
323. ~r2_hidden(sK1,sK0) <- (~1, 21) [subsumption resolution 322,104] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
324. $false <- (~1, 19, 21) [subsumption resolution 323,221] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
325. 1 | ~19 | ~21 [avatar contradiction clause 324]
326. $false [avatar sat refutation 105,110,169,175,222,308,325] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:113,thDist:-113}
% SZS output end Proof for setfam_1__t49_setfam_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 30
% Generated clauses: 160
% Active clauses: 61
% Passive clauses: 66
% Final active clauses: 60
% Final passive clauses: 4
% Main loop iterations started: 62
% 
% >>> Simplifying Inferences
% Duplicate literals: 13
% Trivial inequalities: 3
% Fw subsumption resolutions: 12
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Equational tautologies: 3
% Forward subsumptions: 39
% 
% >>> Generating Inferences
% Binary resolution: 46
% Forward superposition: 11
% Backward superposition: 19
% Self superposition: 1
% Equality factoring: 2
% Equality resolution: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 11
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

