i: [1,41,0,2,1,1,0,0] t36_xxreal_0
a: [3,63,0,2,1,41]
s: 63
i: [1,40,0,2,1,1,0,0] t36_xxreal_0
a: [3,68,0,2,1,40]
s: 68
i: [1,42,0,7,1,1,0,0] t36_xxreal_0
a: [3,58,0,7,1,42]
s: 58
i: [1,54,0,2,1,0,0,1] rc1_xxreal_0
a: [3,88,0,2,1,54]
s: 88
i: [1,46,0,7,3,0,0,255] reflexivity_r1_xxreal_0
a: [3,75,0,5,2,46]
s: 75
i: [1,49,0,9,3,0,0,1] idempotence_k3_xxreal_0
a: [3,79,0,7,2,49]
s: 79
i: [1,50,0,9,3,0,0,1] idempotence_k4_xxreal_0
a: [3,83,0,7,2,50]
s: 83
i: [1,47,0,8,3,0,0,1] fc5_xxreal_0
s: 47
i: [1,48,0,8,3,0,0,1] fc6_xxreal_0
s: 48
i: [1,43,0,9,3,0,0,255] t17_xxreal_0
s: 43
i: [1,53,0,10,4,0,0,255] connectedness_r1_xxreal_0
s: 53
i: [1,51,0,11,3,0,0,1] commutativity_k3_xxreal_0
s: 51
i: [1,52,0,11,3,0,0,1] commutativity_k4_xxreal_0
s: 52
d: [2,106,1,13,3,1,44,58,52]
d: [2,112,1,11,2,2,21,106,63]
a: [3,116,1,4,1,112]
s: 116
d: [2,124,2,0,0,3,52,68,63,116,47]
e: 124
d: [2,125,2,4,2,1,40,116,47]
d: [2,127,2,2,1,2,21,125,63]
d: [2,128,2,0,0,3,21,127,68]
e: 128
r: 116
i: [1,44,0,12,4,0,0,255] d10_xxreal_0
s: 44
i: [1,45,0,12,4,0,0,255] d10_xxreal_0
s: 45
a: [3,115,1,4,1,112]
s: 115
d: [2,132,1,16,4,1,44,58,44]
d: [2,141,1,14,3,2,21,132,63]
a: [3,144,1,5,1,141]
s: 144
a: [3,149,1,5,1,141]
s: 149
a: [3,120,1,7,1,112]
s: 120
d: [2,100,1,13,5,0,44,43,51]
d: [2,103,1,9,3,0,18,100]
s: 103
d: [2,170,2,14,4,1,44,120,44]
d: [2,171,2,11,3,1,19,170]
d: [2,172,2,7,2,2,21,171,115]
d: [2,173,2,5,1,3,21,172,63]
a: [3,177,2,5,1,173]
s: 177
d: [2,185,3,0,0,3,52,63,68,177,103]
e: 185
d: [2,186,3,4,2,1,40,177,103]
d: [2,189,3,2,1,2,21,186,68]
d: [2,190,3,0,0,3,21,189,63]
e: 190
d: [2,108,1,13,3,1,44,58,52]
d: [2,122,1,11,2,2,21,108,63]
f: 190,173,128,122,40,41,42
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_0__t36_xxreal_0
% SZS output start Proof for xxreal_0__t36_xxreal_0
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => k4_xxreal_0(k3_xxreal_0(X0,X1),X1) = X1)) [input t36_xxreal_0]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => k4_xxreal_0(k3_xxreal_0(X0,X1),X1) = X1)) [negated conjecture 1]
3. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => k3_xxreal_0(X0,X1) = k3_xxreal_0(X1,X0)) [input commutativity_k3_xxreal_0]
4. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => k4_xxreal_0(X0,X1) = k4_xxreal_0(X1,X0)) [input commutativity_k4_xxreal_0]
6. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((~r1_xxreal_0(X1,X0) => k4_xxreal_0(X0,X1) = X1) & (r1_xxreal_0(X1,X0) => k4_xxreal_0(X0,X1) = X0)))) [input d10_xxreal_0]
9. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v1_xxreal_0(k3_xxreal_0(X0,X1))) [input fc5_xxreal_0]
15. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => r1_xxreal_0(k3_xxreal_0(X0,X1),X0))) [input t17_xxreal_0]
16. ? [X0] : (? [X1] : (k4_xxreal_0(k3_xxreal_0(X0,X1),X1) != X1 & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
17. ! [X0] : (! [X1] : (r1_xxreal_0(k3_xxreal_0(X0,X1),X0) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 15]
18. ! [X0] : (! [X1] : (((k4_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X1,X0)) & (k4_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X1,X0))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 6]
21. ! [X0,X1] : (v1_xxreal_0(k3_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 9]
22. ! [X0,X1] : (v1_xxreal_0(k3_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 21]
29. ! [X0,X1] : (k3_xxreal_0(X0,X1) = k3_xxreal_0(X1,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 3]
30. ! [X0,X1] : (k3_xxreal_0(X0,X1) = k3_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 29]
31. ! [X0,X1] : (k4_xxreal_0(X0,X1) = k4_xxreal_0(X1,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 4]
32. ! [X0,X1] : (k4_xxreal_0(X0,X1) = k4_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 31]
35. ? [X0] : (? [X1] : (k4_xxreal_0(k3_xxreal_0(X0,X1),X1) != X1 & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (k4_xxreal_0(k3_xxreal_0(sK0,X1),X1) != X1 & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
36. ? [X1] : (k4_xxreal_0(k3_xxreal_0(sK0,X1),X1) != X1 & v1_xxreal_0(X1)) => (sK1 != k4_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) & v1_xxreal_0(sK1)) [choice axiom]
37. (sK1 != k4_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 16,36,35]
40. v1_xxreal_0(sK0) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. v1_xxreal_0(sK1) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. sK1 != k4_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) [cnf transformation 37] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
43. r1_xxreal_0(k3_xxreal_0(X0,X1),X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 17] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. k4_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 18] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
47. v1_xxreal_0(k3_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 22] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k3_xxreal_0(X0,X1) = k3_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 30] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. k4_xxreal_0(X0,X1) = k4_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 32] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. 1 <=> sK1 = k4_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) [avatar definition]
58. sK1 != k4_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) <- (~1) [avatar component clause 56] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
59. ~1 [avatar split clause 42,56]
61. 2 <=> v1_xxreal_0(sK1) [avatar definition]
63. v1_xxreal_0(sK1) <- (2) [avatar component clause 61] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
64. 2 [avatar split clause 41,61]
66. 3 <=> v1_xxreal_0(sK0) [avatar definition]
68. v1_xxreal_0(sK0) <- (3) [avatar component clause 66] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
69. 3 [avatar split clause 40,66]
100. r1_xxreal_0(k3_xxreal_0(X2,X1),X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X2) [superposition 43,51] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
103. r1_xxreal_0(k3_xxreal_0(X2,X1),X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) [duplicate literal removal 100] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
108. sK1 != k4_xxreal_0(sK1,k3_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(k3_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(sK1) <- (~1) [superposition 58,52] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
114. 9 <=> v1_xxreal_0(k3_xxreal_0(sK0,sK1)) [avatar definition]
115. v1_xxreal_0(k3_xxreal_0(sK0,sK1)) <- (9) [avatar component clause 114] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
116. ~v1_xxreal_0(k3_xxreal_0(sK0,sK1)) <- (~9) [avatar component clause 114] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
118. 10 <=> sK1 = k4_xxreal_0(sK1,k3_xxreal_0(sK0,sK1)) [avatar definition]
120. sK1 != k4_xxreal_0(sK1,k3_xxreal_0(sK0,sK1)) <- (~10) [avatar component clause 118] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
122. sK1 != k4_xxreal_0(sK1,k3_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(k3_xxreal_0(sK0,sK1)) <- (~1, 2) [subsumption resolution 108,63] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
123. ~9 | ~10 | 1 | ~2 [avatar split clause 122,61,56,118,114]
125. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~9) [resolution 116,47] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
127. ~v1_xxreal_0(sK0) <- (2, ~9) [subsumption resolution 125,63] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
128. $false <- (2, 3, ~9) [subsumption resolution 127,68] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
129. ~2 | ~3 | 9 [avatar contradiction clause 128]
170. sK1 != sK1 | ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) | ~v1_xxreal_0(k3_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(sK1) <- (~10) [superposition 120,44] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
171. ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) | ~v1_xxreal_0(k3_xxreal_0(sK0,sK1)) | ~v1_xxreal_0(sK1) <- (~10) [trivial inequality removal 170] {a:2,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
172. ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) | ~v1_xxreal_0(sK1) <- (9, ~10) [subsumption resolution 171,115] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
173. ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) <- (2, 9, ~10) [subsumption resolution 172,63] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
175. 13 <=> r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) [avatar definition]
177. ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK1) <- (~13) [avatar component clause 175] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
178. ~13 | ~2 | ~9 | 10 [avatar split clause 173,118,114,61,175]
186. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~13) [resolution 177,103] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
189. ~v1_xxreal_0(sK1) <- (3, ~13) [subsumption resolution 186,68] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
190. $false <- (2, 3, ~13) [subsumption resolution 189,63] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
191. ~2 | ~3 | 13 [avatar contradiction clause 190]
192. $false [avatar sat refutation 59,64,69,123,129,178,191] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:102,thDist:-102}
% SZS output end Proof for xxreal_0__t36_xxreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 102
% Active clauses: 22
% Passive clauses: 32
% Final active clauses: 21
% Final passive clauses: 9
% Main loop iterations started: 22
% 
% >>> Simplifying Inferences
% Duplicate literals: 49
% Trivial inequalities: 2
% Fw subsumption resolutions: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 13
% Backward superposition: 17
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 15
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 4
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% non unit clause index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

