i: [1,58,0,2,1,1,0,0] t41_funct_1
a: [3,118,0,2,1,58]
s: 118
i: [1,57,0,2,1,1,0,0] t41_funct_1
a: [3,123,0,2,1,57]
s: 123
i: [1,56,0,2,1,1,0,0] t41_funct_1
a: [3,128,0,2,1,56]
s: 128
i: [1,55,0,2,1,1,0,0] t41_funct_1
a: [3,133,0,2,1,55]
s: 133
i: [1,54,0,2,1,1,0,0] t41_funct_1
a: [3,138,0,2,1,54]
s: 138
i: [1,61,0,4,1,1,0,0] t41_funct_1
a: [3,103,0,4,1,61]
s: 103
i: [1,59,0,5,1,1,0,0] t41_funct_1
a: [3,113,0,5,1,59]
s: 113
i: [1,60,0,7,1,1,0,0] t41_funct_1
a: [3,108,0,7,1,60]
s: 108
i: [1,95,0,2,1,0,0,1] rc1_funct_1
a: [3,143,0,2,1,95]
s: 143
i: [1,94,0,2,1,0,0,1] rc1_funct_1
a: [3,148,0,2,1,94]
s: 148
i: [1,98,0,2,1,0,0,1] rc2_funct_1
a: [3,153,0,2,1,98]
s: 153
i: [1,97,0,2,1,0,0,1] rc2_funct_1
a: [3,158,0,2,1,97]
s: 158
i: [1,96,0,2,1,0,0,1] rc2_funct_1
a: [3,163,0,2,1,96]
s: 163
i: [1,62,0,3,1,0,0,1] dt_k4_relat_1
s: 62
i: [1,65,0,3,1,0,0,1] fc4_funct_1
s: 65
i: [1,66,0,3,1,0,0,1] fc4_funct_1
s: 66
i: [1,67,0,3,1,0,0,1] fc3_funct_1
s: 67
i: [1,68,0,3,1,0,0,1] fc3_funct_1
s: 68
i: [1,69,0,3,1,0,0,255] fc28_relat_1
s: 69
i: [1,70,0,4,1,0,0,255] fc28_relat_1
s: 70
i: [1,71,0,4,1,0,0,255] fc28_relat_1
s: 71
i: [1,82,0,4,1,0,0,1] dt_k3_relat_1
s: 82
i: [1,87,0,4,1,0,0,255] rc3_relat_1
s: 87
i: [1,63,0,5,1,0,0,1] rd2_relat_1
s: 63
i: [1,64,0,5,1,0,0,1] rd1_relat_1
s: 64
i: [1,88,0,5,1,0,0,255] rc3_relat_1
s: 88
i: [1,89,0,5,1,0,0,255] rc3_relat_1
s: 89
i: [1,75,0,7,3,0,0,1] dt_k2_funct_1
s: 75
i: [1,76,0,7,3,0,0,1] dt_k2_funct_1
s: 76
i: [1,72,0,9,4,0,0,1] fc6_funct_1
s: 72
i: [1,73,0,9,4,0,0,1] fc6_funct_1
s: 73
i: [1,74,0,9,4,0,0,1] fc6_funct_1
s: 74
i: [1,90,0,11,4,0,0,255] fc29_relat_1
s: 90
i: [1,92,0,11,4,0,0,255] fc30_relat_1
s: 92
i: [1,77,0,12,4,0,0,1] t33_funct_1
s: 77
i: [1,78,0,12,4,0,0,1] t33_funct_1
s: 78
i: [1,85,0,12,5,0,0,1] fc2_funct_1
s: 85
i: [1,86,0,12,5,0,0,1] fc2_funct_1
s: 86
i: [1,91,0,12,4,0,0,255] fc29_relat_1
s: 91
i: [1,93,0,12,4,0,0,255] fc30_relat_1
s: 93
i: [1,79,0,14,4,0,0,1] t39_funct_1
s: 79
d: [2,167,1,12,4,1,44,91,108]
d: [2,168,1,10,3,2,21,167,138]
d: [2,169,1,8,2,3,21,168,128]
s: 169
i: [1,80,0,14,4,0,0,1] t39_funct_1
s: 80
d: [2,170,1,12,4,1,44,93,108]
d: [2,171,1,10,3,2,21,170,128]
d: [2,172,1,8,2,3,21,171,138]
s: 172
i: [1,83,0,16,7,0,0,1] fc7_funct_1
s: 83
i: [1,84,0,16,7,0,0,1] fc7_funct_1
s: 84
d: [2,173,1,19,7,0,44,93,79]
d: [2,179,1,17,6,0,18,173]
d: [2,180,1,14,5,0,21,179,75]
s: 180
i: [1,99,0,29,9,0,0,1] l46_funct_1
s: 99
d: [2,196,2,14,5,1,44,180,113]
d: [2,199,2,12,4,2,21,196,128]
d: [2,200,2,10,3,3,21,199,123]
a: [3,204,2,2,1,200]
s: 204
d: [2,212,1,29,9,1,44,99,108]
d: [2,219,1,27,8,2,21,212,123]
d: [2,220,1,25,7,3,21,219,128]
d: [2,221,1,23,6,4,21,220,133]
d: [2,222,1,21,5,5,21,221,138]
d: [2,223,1,21,5,6,22,222,113]
s: 223
d: [2,183,1,19,7,0,44,91,80]
d: [2,187,1,17,6,0,18,183]
d: [2,190,1,14,5,0,21,187,75]
s: 190
d: [2,232,2,31,8,6,44,223,80]
d: [2,233,2,24,7,6,19,232]
d: [2,234,2,20,6,7,21,233,103]
d: [2,235,2,18,5,8,21,234,118]
d: [2,236,2,16,4,8,21,235,133]
d: [2,237,2,14,3,8,21,236,138]
a: [3,241,2,3,1,237]
s: 241
d: [2,256,3,0,0,3,52,138,133,241,75]
e: 256
d: [2,257,3,4,2,1,40,241,75]
d: [2,259,3,2,1,2,21,257,133]
d: [2,260,3,0,0,3,21,259,138]
e: 260
r: 241
d: [2,174,1,20,7,0,44,91,79]
d: [2,178,1,18,6,0,18,174]
d: [2,181,1,15,5,0,21,178,75]
s: 181
a: [3,240,2,3,1,237]
s: 240
d: [2,182,1,20,7,0,44,93,80]
d: [2,188,1,18,6,0,18,182]
d: [2,189,1,15,5,0,21,188,75]
s: 189
d: [2,215,1,32,9,0,44,99,63]
d: [2,228,1,29,8,0,21,215,68]
d: [2,229,1,26,7,0,21,228,62]
s: 229
a: [3,245,2,3,1,237]
s: 245
d: [2,290,3,0,0,3,52,138,133,245,76]
e: 290
d: [2,291,3,4,2,1,40,245,76]
d: [2,293,3,2,1,2,21,291,133]
d: [2,294,3,0,0,3,21,293,138]
e: 294
r: 245
d: [2,214,1,38,12,0,44,99,80]
d: [2,217,1,34,10,0,18,214]
d: [2,226,1,31,9,0,21,217,76]
d: [2,227,1,28,8,0,21,226,75]
s: 227
a: [3,244,2,3,1,237]
s: 244
d: [2,213,1,39,12,0,44,99,79]
d: [2,218,1,35,10,0,18,213]
d: [2,224,1,32,9,0,21,218,76]
d: [2,225,1,29,8,0,21,224,75]
s: 225
d: [2,216,1,39,12,0,44,99,78]
d: [2,230,1,36,11,0,21,216,76]
d: [2,231,1,33,10,0,21,230,75]
s: 231
a: [3,249,2,8,1,237]
s: 249
d: [2,325,3,13,4,1,44,249,78]
d: [2,326,3,6,3,1,19,325]
d: [2,327,3,4,2,2,21,326,118]
d: [2,328,3,2,1,3,21,327,133]
d: [2,329,3,0,0,4,21,328,138]
e: 329
f: 329,294,260,237,54,55,56,57,58,59,60,61
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_1__t41_funct_1
% SZS output start Proof for funct_1__t41_funct_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ((k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(X1) & v2_funct_1(X0)) => k2_funct_1(X0) = X1))) [input t41_funct_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ((k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(X1) & v2_funct_1(X0)) => k2_funct_1(X0) = X1))) [negated conjecture 1]
4. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0)))) [input dt_k2_funct_1]
16. ! [X0] : ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_funct_1(X3) & v1_relat_1(X3)) => ((k4_relat_1(X0) = k3_relat_1(X2,X3) & k3_relat_1(X1,X2) = k4_relat_1(k9_xtuple_0(X3)) & k10_xtuple_0(X1) = X0) => X1 = X3)))) [input l46_funct_1]
22. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))))) [input t33_funct_1]
23. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k3_relat_1(k2_funct_1(X0),X0) = k4_relat_1(k10_xtuple_0(X0)) & k4_relat_1(k9_xtuple_0(X0)) = k3_relat_1(X0,k2_funct_1(X0))))) [input t39_funct_1]
24. ! [X0,X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ! [X3] : ((v1_funct_1(X3) & v1_relat_1(X3)) => ((k4_relat_1(X0) = k3_relat_1(X2,X3) & k3_relat_1(X1,X2) = k4_relat_1(k9_xtuple_0(X3)) & k10_xtuple_0(X1) = X0) => X1 = X3)))) [flattening 16]
25. ? [X0] : (? [X1] : ((k2_funct_1(X0) != X1 & (k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(X1) & v2_funct_1(X0))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
26. ? [X0] : (? [X1] : (k2_funct_1(X0) != X1 & k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(X1) & v2_funct_1(X0) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 25]
29. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 4]
30. ! [X0] : ((v1_funct_1(k2_funct_1(X0)) & v1_relat_1(k2_funct_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 29]
31. ! [X0] : (((k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 22]
32. ! [X0] : ((k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 31]
33. ! [X0] : (((k3_relat_1(k2_funct_1(X0),X0) = k4_relat_1(k10_xtuple_0(X0)) & k4_relat_1(k9_xtuple_0(X0)) = k3_relat_1(X0,k2_funct_1(X0))) | ~v2_funct_1(X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 23]
34. ! [X0] : ((k3_relat_1(k2_funct_1(X0),X0) = k4_relat_1(k10_xtuple_0(X0)) & k4_relat_1(k9_xtuple_0(X0)) = k3_relat_1(X0,k2_funct_1(X0))) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 33]
35. ! [X0,X1] : (! [X2] : (! [X3] : ((X1 = X3 | (k4_relat_1(X0) != k3_relat_1(X2,X3) | k3_relat_1(X1,X2) != k4_relat_1(k9_xtuple_0(X3)) | k10_xtuple_0(X1) != X0)) | (~v1_funct_1(X3) | ~v1_relat_1(X3))) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) [ennf transformation 24]
36. ! [X0,X1] : (! [X2] : (! [X3] : (X1 = X3 | k4_relat_1(X0) != k3_relat_1(X2,X3) | k3_relat_1(X1,X2) != k4_relat_1(k9_xtuple_0(X3)) | k10_xtuple_0(X1) != X0 | ~v1_funct_1(X3) | ~v1_relat_1(X3)) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) [flattening 35]
45. ? [X0] : (? [X1] : (k2_funct_1(X0) != X1 & k3_relat_1(X0,X1) = k4_relat_1(k9_xtuple_0(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(X1) & v2_funct_1(X0) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (k2_funct_1(sK0) != X1 & k3_relat_1(sK0,X1) = k4_relat_1(k9_xtuple_0(sK0)) & k9_xtuple_0(X1) = k10_xtuple_0(sK0) & v2_funct_1(sK0) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
46. ? [X1] : (k2_funct_1(sK0) != X1 & k3_relat_1(sK0,X1) = k4_relat_1(k9_xtuple_0(sK0)) & k9_xtuple_0(X1) = k10_xtuple_0(sK0) & v2_funct_1(sK0) & v1_funct_1(X1) & v1_relat_1(X1)) => (k2_funct_1(sK0) != sK1 & k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) & k10_xtuple_0(sK0) = k9_xtuple_0(sK1) & v2_funct_1(sK0) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
47. (k2_funct_1(sK0) != sK1 & k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) & k10_xtuple_0(sK0) = k9_xtuple_0(sK1) & v2_funct_1(sK0) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 26,46,45]
54. v1_relat_1(sK0) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. v1_funct_1(sK0) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. v1_relat_1(sK1) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
57. v1_funct_1(sK1) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. v2_funct_1(sK0) [cnf transformation 47] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
59. k10_xtuple_0(sK0) = k9_xtuple_0(sK1) [cnf transformation 47] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
60. k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) [cnf transformation 47] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
61. k2_funct_1(sK0) != sK1 [cnf transformation 47] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
75. v1_relat_1(k2_funct_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
76. v1_funct_1(k2_funct_1(X0)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
78. k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 32] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
80. k3_relat_1(k2_funct_1(X0),X0) = k4_relat_1(k10_xtuple_0(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 34] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
81. X1 = X3 | k4_relat_1(X0) != k3_relat_1(X2,X3) | k3_relat_1(X1,X2) != k4_relat_1(k9_xtuple_0(X3)) | k10_xtuple_0(X1) != X0 | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [cnf transformation 36] {a:0,w:32,wCS:12800,sine:1,thAx:0,allAx:1,thDist:-1}
99. k3_relat_1(X2,X3) != k4_relat_1(k10_xtuple_0(X1)) | X1 = X3 | k3_relat_1(X1,X2) != k4_relat_1(k9_xtuple_0(X3)) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [equality resolution 81] {a:0,w:29,wCS:11600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
101. 1 <=> k2_funct_1(sK0) = sK1 [avatar definition]
103. k2_funct_1(sK0) != sK1 <- (~1) [avatar component clause 101] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
104. ~1 [avatar split clause 61,101]
106. 2 <=> k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) [avatar definition]
108. k4_relat_1(k9_xtuple_0(sK0)) = k3_relat_1(sK0,sK1) <- (2) [avatar component clause 106] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
109. 2 [avatar split clause 60,106]
111. 3 <=> k10_xtuple_0(sK0) = k9_xtuple_0(sK1) [avatar definition]
113. k10_xtuple_0(sK0) = k9_xtuple_0(sK1) <- (3) [avatar component clause 111] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 3 [avatar split clause 59,111]
116. 4 <=> v2_funct_1(sK0) [avatar definition]
118. v2_funct_1(sK0) <- (4) [avatar component clause 116] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. 4 [avatar split clause 58,116]
121. 5 <=> v1_funct_1(sK1) [avatar definition]
123. v1_funct_1(sK1) <- (5) [avatar component clause 121] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. 5 [avatar split clause 57,121]
126. 6 <=> v1_relat_1(sK1) [avatar definition]
128. v1_relat_1(sK1) <- (6) [avatar component clause 126] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. 6 [avatar split clause 56,126]
131. 7 <=> v1_funct_1(sK0) [avatar definition]
133. v1_funct_1(sK0) <- (7) [avatar component clause 131] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. 7 [avatar split clause 55,131]
136. 8 <=> v1_relat_1(sK0) [avatar definition]
138. v1_relat_1(sK0) <- (8) [avatar component clause 136] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
139. 8 [avatar split clause 54,136]
212. k4_relat_1(k10_xtuple_0(X0)) != k4_relat_1(k9_xtuple_0(sK0)) | sK1 = X0 | k4_relat_1(k9_xtuple_0(sK1)) != k3_relat_1(X0,sK0) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2) [superposition 99,108] {a:1,w:29,wCS:2900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
219. k4_relat_1(k10_xtuple_0(X0)) != k4_relat_1(k9_xtuple_0(sK0)) | sK1 = X0 | k4_relat_1(k9_xtuple_0(sK1)) != k3_relat_1(X0,sK0) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2, 5) [subsumption resolution 212,123] {a:1,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
220. k4_relat_1(k10_xtuple_0(X0)) != k4_relat_1(k9_xtuple_0(sK0)) | sK1 = X0 | k4_relat_1(k9_xtuple_0(sK1)) != k3_relat_1(X0,sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2, 5, 6) [subsumption resolution 219,128] {a:1,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
221. k4_relat_1(k10_xtuple_0(X0)) != k4_relat_1(k9_xtuple_0(sK0)) | sK1 = X0 | k4_relat_1(k9_xtuple_0(sK1)) != k3_relat_1(X0,sK0) | ~v1_relat_1(sK0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2, 5, 6, 7) [subsumption resolution 220,133] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
222. k4_relat_1(k10_xtuple_0(X0)) != k4_relat_1(k9_xtuple_0(sK0)) | sK1 = X0 | k4_relat_1(k9_xtuple_0(sK1)) != k3_relat_1(X0,sK0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2, 5, 6, 7, 8) [subsumption resolution 221,138] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
223. k4_relat_1(k10_xtuple_0(sK0)) != k3_relat_1(X0,sK0) | k4_relat_1(k10_xtuple_0(X0)) != k4_relat_1(k9_xtuple_0(sK0)) | sK1 = X0 | ~v1_funct_1(X0) | ~v1_relat_1(X0) <- (2, 3, 5, 6, 7, 8) [forward demodulation 222,113] {a:1,w:21,wCS:2100,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
232. k4_relat_1(k10_xtuple_0(sK0)) != k4_relat_1(k10_xtuple_0(sK0)) | k4_relat_1(k9_xtuple_0(sK0)) != k4_relat_1(k10_xtuple_0(k2_funct_1(sK0))) | k2_funct_1(sK0) = sK1 | ~v1_funct_1(k2_funct_1(sK0)) | ~v1_relat_1(k2_funct_1(sK0)) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3, 5, 6, 7, 8) [superposition 223,80] {a:2,w:31,wCS:3100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
233. k4_relat_1(k9_xtuple_0(sK0)) != k4_relat_1(k10_xtuple_0(k2_funct_1(sK0))) | k2_funct_1(sK0) = sK1 | ~v1_funct_1(k2_funct_1(sK0)) | ~v1_relat_1(k2_funct_1(sK0)) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (2, 3, 5, 6, 7, 8) [trivial inequality removal 232] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
234. k4_relat_1(k9_xtuple_0(sK0)) != k4_relat_1(k10_xtuple_0(k2_funct_1(sK0))) | ~v1_funct_1(k2_funct_1(sK0)) | ~v1_relat_1(k2_funct_1(sK0)) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 5, 6, 7, 8) [subsumption resolution 233,103] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
235. k4_relat_1(k9_xtuple_0(sK0)) != k4_relat_1(k10_xtuple_0(k2_funct_1(sK0))) | ~v1_funct_1(k2_funct_1(sK0)) | ~v1_relat_1(k2_funct_1(sK0)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 4, 5, 6, 7, 8) [subsumption resolution 234,118] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
236. k4_relat_1(k9_xtuple_0(sK0)) != k4_relat_1(k10_xtuple_0(k2_funct_1(sK0))) | ~v1_funct_1(k2_funct_1(sK0)) | ~v1_relat_1(k2_funct_1(sK0)) | ~v1_relat_1(sK0) <- (~1, 2, 3, 4, 5, 6, 7, 8) [subsumption resolution 235,133] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
237. k4_relat_1(k9_xtuple_0(sK0)) != k4_relat_1(k10_xtuple_0(k2_funct_1(sK0))) | ~v1_funct_1(k2_funct_1(sK0)) | ~v1_relat_1(k2_funct_1(sK0)) <- (~1, 2, 3, 4, 5, 6, 7, 8) [subsumption resolution 236,138] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
239. 16 <=> v1_relat_1(k2_funct_1(sK0)) [avatar definition]
241. ~v1_relat_1(k2_funct_1(sK0)) <- (~16) [avatar component clause 239] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
243. 17 <=> v1_funct_1(k2_funct_1(sK0)) [avatar definition]
245. ~v1_funct_1(k2_funct_1(sK0)) <- (~17) [avatar component clause 243] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
247. 18 <=> k4_relat_1(k9_xtuple_0(sK0)) = k4_relat_1(k10_xtuple_0(k2_funct_1(sK0))) [avatar definition]
249. k4_relat_1(k9_xtuple_0(sK0)) != k4_relat_1(k10_xtuple_0(k2_funct_1(sK0))) <- (~18) [avatar component clause 247] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
250. ~16 | ~17 | ~18 | 1 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 [avatar split clause 237,136,131,126,121,116,111,106,101,247,243,239]
257. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~16) [resolution 241,75] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
259. ~v1_relat_1(sK0) <- (7, ~16) [subsumption resolution 257,133] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
260. $false <- (7, 8, ~16) [subsumption resolution 259,138] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
261. ~7 | ~8 | 16 [avatar contradiction clause 260]
291. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~17) [resolution 245,76] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
293. ~v1_relat_1(sK0) <- (7, ~17) [subsumption resolution 291,133] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
294. $false <- (7, 8, ~17) [subsumption resolution 293,138] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
295. ~7 | ~8 | 17 [avatar contradiction clause 294]
325. k4_relat_1(k9_xtuple_0(sK0)) != k4_relat_1(k9_xtuple_0(sK0)) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~18) [superposition 249,78] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
326. ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~18) [trivial inequality removal 325] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
327. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (4, ~18) [subsumption resolution 326,118] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
328. ~v1_relat_1(sK0) <- (4, 7, ~18) [subsumption resolution 327,133] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
329. $false <- (4, 7, 8, ~18) [subsumption resolution 328,138] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
330. ~4 | ~7 | ~8 | 18 [avatar contradiction clause 329]
331. $false [avatar sat refutation 104,109,114,119,124,129,134,139,250,261,295,330] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:115,thDist:-115}
% SZS output end Proof for funct_1__t41_funct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 45
% Generated clauses: 202
% Active clauses: 53
% Passive clauses: 78
% Final active clauses: 51
% Final passive clauses: 16
% Main loop iterations started: 62
% 
% >>> Simplifying Inferences
% Duplicate literals: 28
% Trivial inequalities: 2
% Fw subsumption resolutions: 75
% Fw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 19
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 2
% Forward superposition: 34
% Backward superposition: 8
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 5
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 22
% SAT solver unit clauses: 13
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

