i: [1,192,0,2,1,1,0,0] t1_hausdorf
a: [3,290,0,2,1,192]
s: 290
i: [1,191,0,2,1,1,0,0] t1_hausdorf
a: [3,295,0,2,1,191]
s: 295
i: [1,195,0,3,1,1,0,0] t1_hausdorf
a: [3,275,0,3,1,195]
s: 275
i: [1,193,0,3,1,1,0,0] t1_hausdorf
a: [3,285,0,3,1,193]
s: 285
i: [1,194,0,5,1,1,0,0] t1_hausdorf
a: [3,280,0,5,1,194]
s: 280
i: [1,196,0,2,1,0,0,2] fc1_numbers
a: [3,300,0,2,1,196]
s: 300
i: [1,197,0,2,1,0,0,255] fc3_membered
a: [3,305,0,2,1,197]
s: 305
i: [1,198,0,2,1,0,0,255] fc6_membered
a: [3,310,0,2,1,198]
s: 310
i: [1,199,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,315,0,2,1,199]
s: 315
i: [1,262,0,2,1,0,0,1] rc1_xreal_0
a: [3,356,0,2,1,262]
s: 356
i: [1,264,0,2,1,0,0,3] rc2_membered
a: [3,361,0,2,1,264]
s: 361
i: [1,263,0,2,1,0,0,3] rc2_membered
a: [3,366,0,2,1,263]
s: 366
i: [1,266,0,2,1,0,0,3] rc1_membered
a: [3,371,0,2,1,266]
s: 371
i: [1,265,0,2,1,0,0,3] rc1_membered
a: [3,376,0,2,1,265]
s: 376
i: [1,268,0,2,1,0,0,255] rc3_membered
a: [3,381,0,2,1,268]
s: 381
i: [1,267,0,2,1,0,0,255] rc3_membered
a: [3,386,0,2,1,267]
s: 386
i: [1,271,0,2,1,0,0,255] rc4_xreal_0
a: [3,391,0,2,1,271]
s: 391
i: [1,270,0,2,1,0,0,255] rc4_xreal_0
a: [3,396,0,2,1,270]
s: 396
i: [1,269,0,2,1,0,0,255] rc4_xreal_0
a: [3,401,0,2,1,269]
s: 401
i: [1,205,0,3,1,0,0,2] fc1_subset_1
s: 205
i: [1,248,0,3,1,0,0,2] rc2_subset_1
s: 248
i: [1,251,0,3,1,0,0,255] reflexivity_r1_tarski
s: 251
i: [1,200,0,3,1,0,0,255] d2_xboole_0
a: [3,320,0,3,1,200]
s: 320
r: 315
d: [2,403,0,2,1,2,23,315,320]
a: [3,407,0,2,1,403]
s: 407
i: [1,201,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,325,0,3,1,201]
s: 325
r: 310
d: [2,409,0,2,1,2,23,310,325]
a: [3,413,0,2,1,409]
s: 413
i: [1,202,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,330,0,3,1,202]
s: 330
r: 407
r: 320
d: [2,415,0,2,1,2,23,407,330]
a: [3,420,0,2,1,415]
s: 420
d: [2,416,0,3,1,2,23,320,330]
a: [3,425,0,3,1,416]
s: 425
i: [1,212,0,4,2,0,0,255] cc2_xreal_0
s: 212
i: [1,216,0,4,2,0,0,255] cc2_membered
s: 216
i: [1,217,0,4,2,0,0,255] cc3_membered
s: 217
i: [1,223,0,4,2,0,0,255] cc4_membered
s: 223
d: [2,427,1,2,1,1,40,223,305]
a: [3,431,1,2,1,427]
s: 431
i: [1,224,0,4,2,0,0,255] cc5_membered
s: 224
i: [1,227,0,4,2,0,0,255] cc1_membered
s: 227
d: [2,433,1,2,1,1,40,224,305]
a: [3,437,1,2,1,433]
s: 437
i: [1,230,0,4,2,0,0,1] cc4_xreal_0
s: 230
d: [2,463,1,2,1,1,40,230,295]
a: [3,470,1,2,1,463]
s: 470
i: [1,231,0,4,2,0,0,3] cc18_membered
s: 231
i: [1,246,0,4,1,0,0,3] existence_m1_subset_1
s: 246
d: [2,464,1,2,1,1,40,230,290]
a: [3,475,1,2,1,464]
s: 475
i: [1,250,0,4,1,0,0,255] rc3_subset_1
s: 250
d: [2,439,1,2,1,1,40,227,413]
a: [3,446,1,2,1,439]
s: 446
i: [1,203,0,4,1,0,0,2] dt_k5_numbers
a: [3,335,0,4,1,203]
s: 335
i: [1,204,0,4,1,0,0,1] dt_k6_numbers
a: [3,340,0,4,1,204]
s: 340
d: [2,440,1,2,1,1,40,227,361]
a: [3,451,1,2,1,440]
s: 451
i: [1,209,0,5,2,0,0,2] rc1_subset_1
s: 209
d: [2,441,1,2,1,1,40,227,371]
a: [3,456,1,2,1,441]
s: 456
i: [1,232,0,5,2,0,0,2] t6_boole
s: 232
i: [1,235,0,5,2,0,0,2] cc25_membered
s: 235
d: [2,442,1,2,1,1,40,227,381]
a: [3,461,1,2,1,442]
s: 461
i: [1,236,0,5,2,0,0,2] cc1_xreal_0
s: 236
d: [2,465,1,2,1,1,40,230,356]
a: [3,480,1,2,1,465]
s: 480
i: [1,243,0,5,2,0,0,255] t7_boole
s: 243
i: [1,247,0,5,1,0,0,2] rc2_subset_1
s: 247
d: [2,483,1,2,1,1,40,231,401]
a: [3,488,1,2,1,483]
s: 488
i: [1,249,0,5,1,0,0,255] rc3_subset_1
s: 249
d: [2,484,1,2,1,1,40,231,420]
a: [3,493,1,2,1,484]
s: 493
i: [1,257,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,347,0,5,2,257]
s: 347
d: [2,513,0,5,2,1,22,232,330]
s: 513
d: [2,548,1,5,2,2,44,285,513]
s: 548
i: [1,211,0,6,2,0,0,255] rc4_subset_1
s: 211
d: [2,495,2,2,1,1,40,446,216]
a: [3,499,2,2,1,495]
s: 499
i: [1,237,0,6,2,0,0,2] cc11_membered
s: 237
i: [1,238,0,6,2,0,0,255] cc8_membered
s: 238
d: [2,501,2,2,1,1,40,451,216]
a: [3,505,2,2,1,501]
s: 505
i: [1,239,0,6,2,0,0,255] t1_subset
s: 239
d: [2,507,2,2,1,1,40,456,216]
a: [3,511,2,2,1,507]
s: 511
i: [1,252,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 252
i: [1,208,0,7,2,0,0,2] rc1_subset_1
s: 208
d: [2,520,2,2,1,1,40,461,216]
a: [3,524,2,2,1,520]
s: 524
i: [1,210,0,7,2,0,0,255] rc4_subset_1
s: 210
d: [2,605,1,5,2,0,40,208,237]
a: [3,614,1,2,1,605]
s: 614
i: [1,213,0,7,3,0,0,255] cc13_membered
s: 213
i: [1,225,0,7,3,0,0,255] cc14_membered
s: 225
d: [2,532,2,2,1,1,40,488,227]
a: [3,536,2,2,1,532]
s: 536
i: [1,228,0,7,3,0,0,255] cc17_membered
s: 228
d: [2,538,2,2,1,1,40,493,227]
a: [3,542,2,2,1,538]
s: 542
i: [1,241,0,7,2,0,0,255] t3_subset
s: 241
i: [1,242,0,7,3,0,0,1] t8_boole
s: 242
d: [2,549,3,2,1,1,40,499,217]
a: [3,553,3,2,1,549]
s: 553
i: [1,258,0,9,3,0,0,1] idempotence_k4_xxreal_0
a: [3,351,0,7,2,258]
s: 351
d: [2,593,3,2,1,1,40,505,217]
a: [3,597,3,2,1,593]
s: 597
i: [1,214,0,8,3,0,0,255] cc19_membered
s: 214
i: [1,215,0,8,3,0,0,255] cc20_membered
s: 215
d: [2,599,3,2,1,1,40,511,217]
a: [3,603,3,2,1,599]
s: 603
i: [1,218,0,8,3,0,0,255] cc23_membered
s: 218
d: [2,622,3,2,1,1,40,524,217]
a: [3,626,3,2,1,622]
s: 626
i: [1,219,0,8,3,0,0,255] cc22_membered
s: 219
i: [1,226,0,8,3,0,0,255] cc21_membered
s: 226
d: [2,672,3,2,1,1,40,536,216]
a: [3,676,3,2,1,672]
s: 676
i: [1,229,0,8,3,0,0,3] cc24_membered
s: 229
d: [2,693,3,2,1,1,40,542,216]
a: [3,698,3,2,1,693]
s: 698
i: [1,233,0,8,3,0,0,2] cc1_subset_1
s: 233
i: [1,240,0,8,3,0,0,255] t2_subset
s: 240
d: [2,705,4,2,1,1,40,553,224]
a: [3,710,4,2,1,705]
s: 710
i: [1,261,0,8,3,0,0,1] fc36_xreal_0
s: 261
d: [2,838,1,6,3,1,44,261,280]
d: [2,841,1,4,2,2,21,838,290]
d: [2,842,1,2,1,3,21,841,295]
a: [3,846,1,2,1,842]
s: 846
i: [1,220,0,9,3,0,0,1] t25_xxreal_0
s: 220
i: [1,234,0,9,3,0,0,255] cc4_subset_1
s: 234
d: [2,848,2,2,1,1,40,846,230]
a: [3,853,2,2,1,848]
s: 853
i: [1,245,0,9,3,0,0,255] t5_subset
s: 245
d: [2,855,1,7,3,1,44,220,280]
d: [2,858,1,5,2,2,21,855,475]
d: [2,859,1,3,1,3,21,858,470]
a: [3,863,1,3,1,859]
s: 863
i: [1,244,0,10,3,0,0,255] t4_subset
s: 244
i: [1,260,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 260
d: [2,849,2,4,2,2,44,846,513]
s: 849
r: 614
i: [1,206,0,11,4,0,0,255] cc3_subset_1
s: 206
d: [2,892,3,3,1,2,40,849,248]
s: 892
i: [1,207,0,11,4,0,0,255] cc2_subset_1
s: 207
i: [1,259,0,11,3,0,0,1] commutativity_k4_xxreal_0
s: 259
d: [2,912,4,3,1,2,40,892,230]
s: 912
i: [1,222,0,13,5,0,0,1] t1_xxreal_0
s: 222
d: [2,919,3,0,0,5,52,470,853,863,275,285,222]
e: 919
d: [2,920,3,0,0,5,52,853,470,285,275,863,222]
e: 920
d: [2,921,1,10,4,1,40,222,285]
d: [2,932,1,7,3,2,21,921,275]
d: [2,933,1,4,2,3,21,932,863]
d: [2,934,1,2,1,4,21,933,853]
d: [2,935,1,0,0,5,21,934,470]
e: 935
d: [2,926,2,10,4,1,40,222,863]
d: [2,938,2,7,3,2,21,926,275]
d: [2,939,2,4,2,3,21,938,285]
d: [2,940,2,2,1,4,21,939,470]
d: [2,941,2,0,0,5,21,940,853]
e: 941
f: 941,859,848,842,464,463,191,192,193,194,195
% Refutation found. Thanks to Tanya!
% SZS status Theorem for hausdorf__t1_hausdorf
% SZS output start Proof for hausdorf__t1_hausdorf
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((k6_numbers = k4_xxreal_0(X0,X1) & r1_xxreal_0(k6_numbers,X0)) => k6_numbers = X0))) [input t1_hausdorf]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((k6_numbers = k4_xxreal_0(X0,X1) & r1_xxreal_0(k6_numbers,X0)) => k6_numbers = X0))) [negated conjecture 1]
31. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
51. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k4_xxreal_0(X0,X1))) [input fc36_xreal_0]
74. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input t1_xxreal_0]
75. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => r1_xxreal_0(X0,k4_xxreal_0(X0,X1)))) [input t25_xxreal_0]
102. ? [X0] : (? [X1] : ((k6_numbers != X0 & (k6_numbers = k4_xxreal_0(X0,X1) & r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
103. ? [X0] : (? [X1] : (k6_numbers != X0 & k6_numbers = k4_xxreal_0(X0,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 102]
118. ! [X0] : (! [X1] : (r1_xxreal_0(X0,k4_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 75]
121. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 74]
122. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 121]
130. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 31]
163. ! [X0,X1] : (v1_xreal_0(k4_xxreal_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 51]
164. ! [X0,X1] : (v1_xreal_0(k4_xxreal_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 163]
165. ? [X0] : (? [X1] : (k6_numbers != X0 & k6_numbers = k4_xxreal_0(X0,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (k6_numbers != sK0 & k6_numbers = k4_xxreal_0(sK0,X1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
166. ? [X1] : (k6_numbers != sK0 & k6_numbers = k4_xxreal_0(sK0,X1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) => (k6_numbers != sK0 & k6_numbers = k4_xxreal_0(sK0,sK1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) [choice axiom]
167. (k6_numbers != sK0 & k6_numbers = k4_xxreal_0(sK0,sK1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 103,166,165]
191. v1_xreal_0(sK0) [cnf transformation 167] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
192. v1_xreal_0(sK1) [cnf transformation 167] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
193. r1_xxreal_0(k6_numbers,sK0) [cnf transformation 167] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
194. k6_numbers = k4_xxreal_0(sK0,sK1) [cnf transformation 167] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
195. k6_numbers != sK0 [cnf transformation 167] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
220. r1_xxreal_0(X0,k4_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 118] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
222. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 122] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
230. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 130] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
261. v1_xreal_0(k4_xxreal_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 164] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
273. 1 <=> k6_numbers = sK0 [avatar definition]
275. k6_numbers != sK0 <- (~1) [avatar component clause 273] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
276. ~1 [avatar split clause 195,273]
278. 2 <=> k6_numbers = k4_xxreal_0(sK0,sK1) [avatar definition]
280. k6_numbers = k4_xxreal_0(sK0,sK1) <- (2) [avatar component clause 278] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
281. 2 [avatar split clause 194,278]
283. 3 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
285. r1_xxreal_0(k6_numbers,sK0) <- (3) [avatar component clause 283] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
286. 3 [avatar split clause 193,283]
288. 4 <=> v1_xreal_0(sK1) [avatar definition]
290. v1_xreal_0(sK1) <- (4) [avatar component clause 288] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
291. 4 [avatar split clause 192,288]
293. 5 <=> v1_xreal_0(sK0) [avatar definition]
295. v1_xreal_0(sK0) <- (5) [avatar component clause 293] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
296. 5 [avatar split clause 191,293]
463. v1_xxreal_0(sK0) <- (5) [resolution 230,295] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
464. v1_xxreal_0(sK1) <- (4) [resolution 230,290] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
468. 38 <=> v1_xxreal_0(sK0) [avatar definition]
470. v1_xxreal_0(sK0) <- (38) [avatar component clause 468] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
471. 38 | ~5 [avatar split clause 463,293,468]
473. 39 <=> v1_xxreal_0(sK1) [avatar definition]
475. v1_xxreal_0(sK1) <- (39) [avatar component clause 473] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
476. 39 | ~4 [avatar split clause 464,288,473]
838. v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (2) [superposition 261,280] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
841. v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (2, 4) [subsumption resolution 838,290] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
842. v1_xreal_0(k6_numbers) <- (2, 4, 5) [subsumption resolution 841,295] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
844. 85 <=> v1_xreal_0(k6_numbers) [avatar definition]
846. v1_xreal_0(k6_numbers) <- (85) [avatar component clause 844] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
847. 85 | ~2 | ~4 | ~5 [avatar split clause 842,293,288,278,844]
848. v1_xxreal_0(k6_numbers) <- (85) [resolution 846,230] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
851. 86 <=> v1_xxreal_0(k6_numbers) [avatar definition]
853. v1_xxreal_0(k6_numbers) <- (86) [avatar component clause 851] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
854. 86 | ~85 [avatar split clause 848,844,851]
855. r1_xxreal_0(sK0,k6_numbers) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (2) [superposition 220,280] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
858. r1_xxreal_0(sK0,k6_numbers) | ~v1_xxreal_0(sK0) <- (2, 39) [subsumption resolution 855,475] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
859. r1_xxreal_0(sK0,k6_numbers) <- (2, 38, 39) [subsumption resolution 858,470] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
861. 87 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
863. r1_xxreal_0(sK0,k6_numbers) <- (87) [avatar component clause 861] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
864. 87 | ~2 | ~38 | ~39 [avatar split clause 859,473,468,278,861]
926. k6_numbers = sK0 | ~r1_xxreal_0(k6_numbers,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k6_numbers) <- (87) [resolution 222,863] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
938. ~r1_xxreal_0(k6_numbers,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k6_numbers) <- (~1, 87) [subsumption resolution 926,275] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
939. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k6_numbers) <- (~1, 3, 87) [subsumption resolution 938,285] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
940. ~v1_xxreal_0(k6_numbers) <- (~1, 3, 38, 87) [subsumption resolution 939,470] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
941. $false <- (~1, 3, 38, 86, 87) [subsumption resolution 940,853] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
942. 1 | ~3 | ~38 | ~86 | ~87 [avatar contradiction clause 941]
943. $false [avatar sat refutation 276,281,286,291,296,471,476,847,854,864,942] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:469,thDist:-469}
% SZS output end Proof for hausdorf__t1_hausdorf
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 81
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 5
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 81
% Generated clauses: 397
% Active clauses: 111
% Passive clauses: 237
% Final active clauses: 106
% Final passive clauses: 124
% Main loop iterations started: 113
% 
% >>> Simplifying Inferences
% Duplicate literals: 15
% Fw subsumption resolutions: 26
% Fw demodulations: 1
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Equational tautologies: 1
% Forward subsumptions: 22
% 
% >>> Generating Inferences
% Binary resolution: 169
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 9
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 16
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 93
% SAT solver unit clauses: 32
% SAT solver binary clauses: 42
% 
% Memory used [KB]: 9978
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% forward subsumption resolution: 0.001 s
% forward demodulation: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

