i: [1,278,0,2,1,1,0,0] l93_xxreal_3
a: [3,742,0,2,1,278]
s: 742
i: [1,277,0,2,1,1,0,0] l93_xxreal_3
a: [3,747,0,2,1,277]
s: 747
i: [1,276,0,2,1,1,0,0] l93_xxreal_3
a: [3,752,0,2,1,276]
s: 752
i: [1,275,0,2,1,1,0,0] l93_xxreal_3
a: [3,757,0,2,1,275]
s: 757
i: [1,280,0,2,1,0,0,255] fc2_xreal_0
a: [3,762,0,2,1,280]
s: 762
i: [1,281,0,2,1,0,0,255] fc1_xreal_0
a: [3,767,0,2,1,281]
s: 767
i: [1,282,0,2,1,0,0,255] fc1_xboole_0
a: [3,772,0,2,1,282]
s: 772
i: [1,283,0,2,1,0,0,255] fc4_xxreal_0
a: [3,777,0,2,1,283]
s: 777
i: [1,284,0,2,1,0,0,255] fc3_xxreal_0
a: [3,782,0,2,1,284]
s: 782
i: [1,285,0,2,1,0,0,255] fc2_xxreal_0
a: [3,787,0,2,1,285]
s: 787
i: [1,286,0,2,1,0,0,255] fc1_xxreal_0
a: [3,792,0,2,1,286]
s: 792
i: [1,295,0,2,1,0,0,255] fc6_ordinal1
a: [3,832,0,2,1,295]
s: 832
i: [1,294,0,2,1,0,0,255] fc6_ordinal1
a: [3,837,0,2,1,294]
s: 837
i: [1,546,0,21,9,0,0,255] d5_xxreal_3
d: [2,733,0,17,7,0,18,546]
a: [3,842,0,2,1,733]
s: 842
i: [1,502,0,2,1,0,0,2] rc2_xboole_0
a: [3,960,0,2,1,502]
s: 960
i: [1,503,0,2,1,0,0,255] rc4_ordinal1
a: [3,965,0,2,1,503]
s: 965
i: [1,504,0,2,1,0,0,255] rc1_ordinal1
a: [3,970,0,2,1,504]
s: 970
i: [1,505,0,2,1,0,0,255] rc1_xcmplx_0
a: [3,975,0,2,1,505]
s: 975
i: [1,506,0,2,1,0,0,1] rc1_xreal_0
a: [3,980,0,2,1,506]
s: 980
i: [1,507,0,2,1,0,0,2] rc1_xboole_0
a: [3,985,0,2,1,507]
s: 985
i: [1,508,0,2,1,0,0,1] rc1_xxreal_0
a: [3,990,0,2,1,508]
s: 990
i: [1,512,0,2,1,0,0,255] rc2_ordinal1
a: [3,995,0,2,1,512]
s: 995
i: [1,511,0,2,1,0,0,255] rc2_ordinal1
a: [3,1000,0,2,1,511]
s: 1000
i: [1,510,0,2,1,0,0,255] rc2_ordinal1
a: [3,1005,0,2,1,510]
s: 1005
i: [1,509,0,2,1,0,0,255] rc2_ordinal1
a: [3,1010,0,2,1,509]
s: 1010
i: [1,514,0,2,1,0,0,255] rc5_ordinal1
a: [3,1015,0,2,1,514]
s: 1015
i: [1,513,0,2,1,0,0,255] rc5_ordinal1
a: [3,1020,0,2,1,513]
s: 1020
i: [1,516,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,1025,0,2,1,516]
s: 1025
i: [1,515,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,1030,0,2,1,515]
s: 1030
i: [1,520,0,2,1,0,0,255] rc3_xreal_0
a: [3,1035,0,2,1,520]
s: 1035
i: [1,519,0,2,1,0,0,255] rc3_xreal_0
a: [3,1040,0,2,1,519]
s: 1040
i: [1,518,0,2,1,0,0,255] rc3_xreal_0
a: [3,1045,0,2,1,518]
s: 1045
i: [1,517,0,2,1,0,0,255] rc3_xreal_0
a: [3,1050,0,2,1,517]
s: 1050
i: [1,524,0,2,1,0,0,255] rc2_xreal_0
a: [3,1055,0,2,1,524]
s: 1055
i: [1,523,0,2,1,0,0,255] rc2_xreal_0
a: [3,1060,0,2,1,523]
s: 1060
i: [1,522,0,2,1,0,0,255] rc2_xreal_0
a: [3,1065,0,2,1,522]
s: 1065
i: [1,521,0,2,1,0,0,255] rc2_xreal_0
a: [3,1070,0,2,1,521]
s: 1070
i: [1,528,0,2,1,0,0,255] rc4_xreal_0
a: [3,1075,0,2,1,528]
s: 1075
i: [1,527,0,2,1,0,0,255] rc4_xreal_0
a: [3,1080,0,2,1,527]
s: 1080
i: [1,526,0,2,1,0,0,255] rc4_xreal_0
a: [3,1085,0,2,1,526]
s: 1085
i: [1,525,0,2,1,0,0,255] rc4_xreal_0
a: [3,1090,0,2,1,525]
s: 1090
i: [1,530,0,2,1,0,0,2] rc4_xxreal_0
a: [3,1095,0,2,1,530]
s: 1095
i: [1,529,0,2,1,0,0,2] rc4_xxreal_0
a: [3,1100,0,2,1,529]
s: 1100
i: [1,532,0,2,1,0,0,2] rc3_xxreal_0
a: [3,1105,0,2,1,532]
s: 1105
i: [1,531,0,2,1,0,0,2] rc3_xxreal_0
a: [3,1110,0,2,1,531]
s: 1110
i: [1,535,0,2,1,0,0,2] rc2_xxreal_3
a: [3,1115,0,2,1,535]
s: 1115
i: [1,534,0,2,1,0,0,2] rc2_xxreal_3
a: [3,1120,0,2,1,534]
s: 1120
i: [1,533,0,2,1,0,0,2] rc2_xxreal_3
a: [3,1125,0,2,1,533]
s: 1125
i: [1,537,0,2,1,0,0,2] rc2_xxreal_0
a: [3,1130,0,2,1,537]
s: 1130
i: [1,536,0,2,1,0,0,2] rc2_xxreal_0
a: [3,1135,0,2,1,536]
s: 1135
i: [1,540,0,2,1,0,0,2] rc1_xxreal_3
a: [3,1140,0,2,1,540]
s: 1140
i: [1,539,0,2,1,0,0,2] rc1_xxreal_3
a: [3,1145,0,2,1,539]
s: 1145
i: [1,538,0,2,1,0,0,2] rc1_xxreal_3
a: [3,1150,0,2,1,538]
s: 1150
i: [1,545,0,2,1,0,0,255] rc1_real_1
a: [3,1155,0,2,1,545]
s: 1155
i: [1,544,0,2,1,0,0,255] rc1_real_1
a: [3,1160,0,2,1,544]
s: 1160
i: [1,543,0,2,1,0,0,255] rc1_real_1
a: [3,1165,0,2,1,543]
s: 1165
i: [1,542,0,2,1,0,0,255] rc1_real_1
a: [3,1170,0,2,1,542]
s: 1170
i: [1,279,0,11,1,1,0,0] l93_xxreal_3
a: [3,737,0,11,1,279]
s: 737
i: [1,468,0,3,1,0,0,255] reflexivity_r1_tarski
s: 468
i: [1,287,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,797,0,3,1,287]
s: 797
i: [1,288,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,802,0,3,1,288]
s: 802
r: 772
d: [2,1177,0,2,1,2,23,772,802]
a: [3,1181,0,2,1,1177]
s: 1181
i: [1,289,0,3,1,0,0,255] d2_xxreal_0
a: [3,807,0,3,1,289]
s: 807
r: 792
r: 782
r: 762
d: [2,1183,0,2,1,2,23,792,807]
a: [3,1189,0,2,1,1183]
s: 1189
d: [2,1184,0,2,1,2,23,782,807]
a: [3,1194,0,2,1,1184]
s: 1194
d: [2,1185,0,2,1,2,23,762,807]
a: [3,1199,0,2,1,1185]
s: 1199
i: [1,290,0,3,1,0,0,255] t1_numerals
a: [3,812,0,3,1,290]
s: 812
i: [1,541,0,3,1,0,0,255] rc1_real_1
a: [3,1175,0,3,1,541]
s: 1175
d: [2,1201,0,3,1,2,22,812,802]
a: [3,1205,0,3,1,1201]
s: 1205
i: [1,297,0,4,2,0,0,255] cc6_ordinal1
s: 297
i: [1,298,0,4,2,0,0,255] cc2_xcmplx_0
s: 298
d: [2,1207,1,2,1,1,40,297,965]
a: [3,1212,1,2,1,1207]
s: 1212
i: [1,299,0,4,2,0,0,255] cc2_xreal_0
s: 299
d: [2,1208,1,2,1,1,40,297,1015]
a: [3,1217,1,2,1,1208]
s: 1217
i: [1,300,0,4,2,0,0,255] cc2_xxreal_0
s: 300
i: [1,301,0,4,2,0,0,255] cc1_ordinal1
s: 301
d: [2,1219,1,2,1,1,40,298,965]
a: [3,1224,1,2,1,1219]
s: 1224
i: [1,302,0,4,2,0,0,255] cc1_ordinal1
s: 302
d: [2,1220,1,2,1,1,40,298,1015]
a: [3,1229,1,2,1,1220]
s: 1229
i: [1,305,0,4,2,0,0,255] cc3_xreal_0
s: 305
i: [1,306,0,4,2,0,0,1] cc4_xreal_0
s: 306
d: [2,1231,1,2,1,1,40,299,965]
a: [3,1236,1,2,1,1231]
s: 1236
i: [1,307,0,4,2,0,0,255] cc9_ordinal1
s: 307
d: [2,1232,1,2,1,1,40,299,1015]
a: [3,1241,1,2,1,1232]
s: 1241
i: [1,308,0,4,2,0,0,255] cc7_ordinal1
s: 308
i: [1,309,0,4,2,0,0,255] cc3_ordinal1
s: 309
d: [2,1243,1,2,1,1,40,300,965]
a: [3,1248,1,2,1,1243]
s: 1248
i: [1,467,0,4,1,0,0,255] existence_m1_subset_1
s: 467
d: [2,1244,1,2,1,1,40,300,1015]
a: [3,1253,1,2,1,1244]
s: 1253
i: [1,291,0,4,1,0,0,255] dt_k5_numbers
a: [3,817,0,4,1,291]
s: 817
i: [1,292,0,4,1,0,0,255] dt_k6_numbers
a: [3,822,0,4,1,292]
s: 822
d: [2,1255,1,2,1,1,40,301,832]
a: [3,1263,1,2,1,1255]
s: 1263
d: [2,1403,0,4,1,2,22,817,797]
a: [3,1407,0,4,1,1403]
s: 1407
d: [2,1257,1,2,1,1,40,301,970]
a: [3,1273,1,2,1,1257]
s: 1273
d: [2,1409,0,4,1,2,22,822,797]
a: [3,1413,0,4,1,1409]
s: 1413
i: [1,310,0,5,2,0,0,255] t6_boole
s: 310
d: [2,1280,1,2,1,1,40,302,832]
a: [3,1288,1,2,1,1280]
s: 1288
i: [1,436,0,5,2,0,0,255] cc8_ordinal1
s: 436
d: [2,1282,1,2,1,1,40,302,970]
a: [3,1298,1,2,1,1282]
s: 1298
i: [1,437,0,5,2,0,0,255] cc1_xcmplx_0
s: 437
i: [1,438,0,5,2,0,0,255] cc1_real_1
s: 438
d: [2,1305,1,2,1,1,40,305,980]
a: [3,1313,1,2,1,1305]
s: 1313
i: [1,439,0,5,2,0,0,255] cc1_xreal_0
s: 439
d: [2,1315,1,2,1,1,40,306,980]
a: [3,1323,1,2,1,1315]
s: 1323
i: [1,464,0,5,2,0,0,255] t7_boole
s: 464
i: [1,293,0,5,1,0,0,255] d3_xxreal_0
a: [3,827,0,5,1,293]
s: 827
d: [2,1328,1,2,1,1,40,307,1181]
a: [3,1335,1,2,1,1328]
s: 1335
d: [2,1415,0,5,2,1,22,310,802]
s: 1415
d: [2,1329,1,2,1,1,40,307,985]
a: [3,1340,1,2,1,1329]
s: 1340
i: [1,444,0,6,3,0,0,255] cc2_ordinal1
s: 444
i: [1,446,0,6,3,0,0,2] cc7_xxreal_0
s: 446
d: [2,1330,1,2,1,1,40,307,1090]
a: [3,1345,1,2,1,1330]
s: 1345
i: [1,447,0,6,3,0,0,2] cc7_xxreal_0
s: 447
d: [2,1331,1,2,1,1,40,307,1100]
a: [3,1350,1,2,1,1331]
s: 1350
i: [1,454,0,6,3,0,0,2] cc5_xxreal_0
s: 454
i: [1,456,0,6,3,0,0,2] cc5_xxreal_0
s: 456
d: [2,1355,1,2,1,1,40,308,1181]
a: [3,1362,1,2,1,1355]
s: 1362
d: [2,1529,2,0,0,2,52,842,1362,300]
e: 1529
d: [2,1530,2,2,1,1,40,1362,300]
d: [2,1536,2,0,0,2,21,1530,842]
e: 1536
r: 842
a: [3,841,0,2,1,733]
s: 841
d: [2,1356,1,2,1,1,40,308,985]
a: [3,1367,1,2,1,1356]
s: 1367
i: [1,457,0,6,3,0,0,2] cc3_xxreal_0
s: 457
i: [1,459,0,6,3,0,0,2] cc3_xxreal_0
s: 459
d: [2,1357,1,2,1,1,40,308,1090]
a: [3,1372,1,2,1,1357]
s: 1372
i: [1,460,0,6,2,0,0,255] t1_subset
s: 460
d: [2,1358,1,2,1,1,40,308,1100]
a: [3,1377,1,2,1,1358]
s: 1377
i: [1,469,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 469
i: [1,303,0,7,3,0,0,255] cc5_ordinal1
s: 303
d: [2,1379,1,2,1,1,40,309,1181]
a: [3,1386,1,2,1,1379]
s: 1386
i: [1,304,0,7,2,0,0,255] t2_arithm
s: 304
d: [2,1380,1,2,1,1,40,309,985]
a: [3,1391,1,2,1,1380]
s: 1391
i: [1,462,0,7,2,0,0,255] t3_subset
s: 462
i: [1,463,0,7,3,0,0,1] t8_boole
s: 463
d: [2,1381,1,2,1,1,40,309,1090]
a: [3,1396,1,2,1,1381]
s: 1396
i: [1,296,0,8,3,0,0,255] cc10_ordinal1
s: 296
d: [2,1382,1,2,1,1,40,309,1100]
a: [3,1401,1,2,1,1382]
s: 1401
i: [1,441,0,8,4,0,0,2] cc4_xxreal_0
s: 441
i: [1,443,0,8,4,0,0,2] cc6_xxreal_0
s: 443
d: [2,1678,1,6,3,1,40,441,757]
a: [3,1697,1,2,1,1678]
s: 1697
r: 1697
i: [1,449,0,8,4,0,0,2] cc1_xxreal_3
s: 449
a: [3,1696,1,2,1,1678]
s: 1696
i: [1,451,0,8,4,0,0,2] cc2_xxreal_3
s: 451
i: [1,452,0,8,4,0,0,2] cc8_xxreal_0
s: 452
a: [3,1701,1,2,1,1678]
s: 1701
i: [1,461,0,8,3,0,0,255] t2_subset
s: 461
a: [3,1704,1,2,1,1678]
s: 1704
i: [1,481,0,8,3,0,0,255] fc3_xcmplx_0
s: 481
i: [1,483,0,8,3,0,0,255] fc6_xreal_0
s: 483
d: [2,1679,1,6,3,1,40,441,752]
a: [3,1710,1,2,1,1679]
s: 1710
r: 1710
i: [1,498,0,8,3,0,0,1] dt_k4_xxreal_3
s: 498
a: [3,1709,1,2,1,1679]
s: 1709
i: [1,466,0,9,3,0,0,255] t5_subset
s: 466
i: [1,465,0,10,3,0,0,255] t4_subset
s: 465
a: [3,1714,1,2,1,1679]
s: 1714
i: [1,484,0,10,4,0,0,1] fc37_xxreal_3
s: 484
a: [3,1717,1,2,1,1679]
s: 1717
i: [1,485,0,10,4,0,0,1] fc37_xxreal_3
s: 485
i: [1,311,0,11,4,0,0,255] t14_xxreal_0
s: 311
d: [2,1680,1,6,3,1,40,441,747]
a: [3,1723,1,2,1,1680]
s: 1723
r: 1723
i: [1,482,0,11,3,0,0,255] commutativity_k3_xcmplx_0
s: 482
a: [3,1722,1,2,1,1680]
s: 1722
i: [1,499,0,11,3,0,0,1] commutativity_k4_xxreal_3
s: 499
d: [2,1927,0,11,4,1,22,311,807]
s: 1927
a: [3,1727,1,2,1,1680]
s: 1727
i: [1,474,0,12,5,0,0,255] fc8_xcmplx_0
s: 474
a: [3,1730,1,2,1,1680]
s: 1730
i: [1,475,0,12,5,0,0,1] fc36_xxreal_3
s: 475
i: [1,476,0,12,5,0,0,1] fc36_xxreal_3
s: 476
d: [2,1973,1,17,3,1,44,737,499]
d: [2,1983,1,15,2,2,21,1973,747]
a: [3,1987,1,4,1,1983]
s: 1987
d: [2,2017,2,0,0,3,52,757,752,1987,498]
e: 2017
d: [2,2018,2,4,2,1,40,1987,498]
d: [2,2020,2,2,1,2,21,2018,752]
d: [2,2021,2,0,0,3,21,2020,757]
e: 2021
r: 1987
i: [1,477,0,12,5,0,0,255] fc26_xreal_0
s: 477
a: [3,1986,1,4,1,1983]
s: 1986
i: [1,478,0,12,5,0,0,255] fc24_xreal_0
s: 478
i: [1,479,0,12,5,0,0,255] fc23_xreal_0
s: 479
d: [2,2026,2,12,3,1,40,1986,449]
a: [3,2031,2,4,1,2026]
s: 2031
i: [1,480,0,12,5,0,0,255] fc25_xreal_0
s: 480
a: [3,2034,2,4,1,2026]
s: 2034
i: [1,486,0,12,5,0,0,1] fc29_xxreal_3
s: 486
i: [1,487,0,12,5,0,0,1] fc29_xxreal_3
s: 487
d: [2,1449,1,4,2,1,40,446,985]
a: [3,1455,1,2,1,1449]
s: 1455
i: [1,488,0,12,5,0,0,1] fc27_xxreal_3
s: 488
a: [3,1458,1,2,1,1449]
s: 1458
i: [1,489,0,12,5,0,0,1] fc27_xxreal_3
s: 489
i: [1,490,0,12,5,0,0,1] fc28_xxreal_3
s: 490
d: [2,1450,1,4,2,1,40,446,1090]
d: [2,1461,1,2,1,2,21,1450,1080]
a: [3,1465,1,2,1,1461]
s: 1465
i: [1,491,0,12,5,0,0,1] fc28_xxreal_3
s: 491
d: [2,1451,1,4,2,1,40,446,1100]
d: [2,1467,1,2,1,2,21,1451,1095]
a: [3,1471,1,2,1,1467]
s: 1471
i: [1,492,0,12,5,0,0,1] fc25_xxreal_3
s: 492
i: [1,493,0,12,5,0,0,1] fc25_xxreal_3
s: 493
d: [2,1474,1,4,2,1,40,447,985]
a: [3,1480,1,2,1,1474]
s: 1480
i: [1,494,0,12,5,0,0,1] fc24_xxreal_3
s: 494
d: [2,1475,1,4,2,1,40,447,1090]
d: [2,1482,1,2,1,2,21,1475,1080]
a: [3,1486,1,2,1,1482]
s: 1486
i: [1,495,0,12,5,0,0,1] fc24_xxreal_3
s: 495
i: [1,496,0,12,5,0,0,1] fc26_xxreal_3
s: 496
d: [2,1476,1,4,2,1,40,447,1100]
d: [2,1488,1,2,1,2,21,1476,1095]
a: [3,1492,1,2,1,1488]
s: 1492
i: [1,497,0,12,5,0,0,1] fc26_xxreal_3
s: 497
d: [2,1494,1,4,2,1,40,456,1194]
d: [2,1499,1,2,1,2,21,1494,1189]
a: [3,1503,1,2,1,1499]
s: 1503
i: [1,473,0,14,4,0,0,255] existence_m2_subset_1
s: 473
i: [1,470,0,15,5,0,0,255] dt_m2_subset_1
s: 470
d: [2,1495,1,4,2,1,40,456,1060]
d: [2,1505,1,2,1,2,21,1495,1065]
a: [3,1509,1,2,1,1505]
s: 1509
i: [1,471,0,15,5,0,0,255] redefinition_m2_subset_1
s: 471
d: [2,1496,1,4,2,1,40,456,1130]
d: [2,1511,1,2,1,2,21,1496,1135]
a: [3,1515,1,2,1,1511]
s: 1515
i: [1,472,0,15,5,0,0,255] redefinition_m2_subset_1
s: 472
i: [1,673,0,15,5,0,0,255] ie4_xxreal_3
s: 673
d: [2,1497,1,4,2,1,40,456,1145]
d: [2,1517,1,2,1,2,21,1497,1150]
a: [3,1521,1,2,1,1517]
s: 1521
d: [2,2132,0,13,4,0,21,673,305]
d: [2,2133,0,11,3,0,21,2132,305]
s: 2133
d: [2,1498,1,4,2,1,40,456,1160]
d: [2,1523,1,2,1,2,21,1498,1165]
a: [3,1527,1,2,1,1523]
s: 1527
a: [3,845,0,15,6,733]
s: 845
i: [1,547,0,21,9,0,0,255] d5_xxreal_3
d: [2,732,0,17,7,0,18,547]
a: [3,849,0,15,6,732]
s: 849
d: [2,1473,1,4,2,1,40,447,1181]
a: [3,1542,1,2,1,1473]
s: 1542
i: [1,582,0,21,9,0,0,255] d5_xxreal_3
d: [2,707,0,17,7,0,18,582]
a: [3,911,0,15,6,707]
s: 911
d: [2,1448,1,4,2,1,40,446,1181]
a: [3,1547,1,2,1,1448]
s: 1547
i: [1,583,0,21,9,0,0,255] d5_xxreal_3
d: [2,706,0,17,7,0,18,583]
a: [3,915,0,15,6,706]
s: 915
i: [1,646,0,17,7,0,0,255] d5_xxreal_3
a: [3,919,0,15,6,646]
s: 919
d: [2,1600,1,4,2,1,40,303,1175]
a: [3,1606,1,2,1,1600]
s: 1606
i: [1,647,0,17,7,0,0,255] d5_xxreal_3
a: [3,923,0,15,6,647]
s: 923
d: [2,1652,1,4,2,1,40,296,1407]
a: [3,1657,1,2,1,1652]
s: 1657
i: [1,648,0,17,7,0,0,255] d5_xxreal_3
a: [3,927,0,15,6,648]
s: 927
i: [1,649,0,17,7,0,0,255] d5_xxreal_3
a: [3,931,0,15,6,649]
s: 931
d: [2,1683,1,6,3,1,40,441,990]
a: [3,1762,1,2,1,1683]
s: 1762
r: 1762
i: [1,654,0,17,7,0,0,255] d5_xxreal_3
a: [3,935,0,15,6,654]
s: 935
a: [3,1761,1,2,1,1683]
s: 1761
i: [1,655,0,17,7,0,0,255] d5_xxreal_3
a: [3,939,0,15,6,655]
s: 939
i: [1,656,0,17,7,0,0,255] d5_xxreal_3
a: [3,943,0,15,6,656]
s: 943
a: [3,1766,1,2,1,1683]
s: 1766
i: [1,657,0,17,7,0,0,255] d5_xxreal_3
a: [3,947,0,15,6,657]
s: 947
a: [3,1769,1,2,1,1683]
s: 1769
i: [1,670,0,17,4,0,0,255] l91_xxreal_3
a: [3,951,0,15,3,670]
s: 951
i: [1,671,0,17,4,0,0,255] l90_xxreal_3
a: [3,955,0,15,3,671]
s: 955
d: [2,1851,1,5,2,1,40,461,1175]
a: [3,1874,1,2,1,1851]
s: 1874
d: [2,2643,2,0,0,2,52,1657,1874,307]
e: 2643
d: [2,2644,2,0,0,2,52,1606,1874,309]
e: 2644
d: [2,2645,2,0,0,3,52,1194,1189,1874,446]
e: 2645
d: [2,2647,2,4,2,1,40,1874,446]
d: [2,2654,2,2,1,2,21,2647,1189]
d: [2,2655,2,0,0,3,21,2654,1194]
e: 2655
d: [2,2648,2,2,1,1,40,1874,309]
d: [2,2657,2,0,0,2,21,2648,1606]
e: 2657
d: [2,2650,2,2,1,1,40,1874,307]
d: [2,2659,2,0,0,2,21,2650,1657]
e: 2659
r: 1874
d: [2,2478,0,15,6,2,22,935,807]
s: 2478
a: [3,1873,1,2,1,1851]
s: 1873
d: [2,2479,0,15,6,2,22,939,807]
s: 2479
d: [2,2480,0,15,6,2,22,943,807]
s: 2480
d: [2,2826,1,21,6,3,44,737,2480]
d: [2,2865,1,17,5,4,21,2826,1986]
d: [2,2866,1,15,4,5,21,2865,747]
d: [2,2867,1,11,3,6,21,2866,2031]
d: [2,2868,1,9,2,7,21,2867,1727]
a: [3,2872,1,7,1,2868]
s: 2872
d: [2,2484,0,15,6,2,22,947,807]
s: 2484
d: [2,2892,2,19,6,2,44,2872,849]
d: [2,2913,2,17,5,3,21,2892,757]
d: [2,2914,2,15,4,4,21,2913,742]
a: [3,2926,2,4,1,2914]
s: 2926
d: [2,3035,3,0,0,3,52,752,747,2926,498]
e: 3035
d: [2,3036,3,4,2,1,40,2926,498]
d: [2,3048,3,2,1,2,21,3036,747]
d: [2,3049,3,0,0,3,21,3048,752]
e: 3049
r: 2926
i: [1,500,0,17,4,0,0,255] spc7_arithm
s: 500
i: [1,548,0,21,9,0,0,255] d5_xxreal_3
d: [2,731,0,19,8,0,18,548]
a: [3,853,0,17,7,731]
s: 853
a: [3,2930,2,3,1,2914]
s: 2930
i: [1,549,0,21,9,0,0,255] d5_xxreal_3
d: [2,730,0,19,8,0,18,549]
a: [3,857,0,17,7,730]
s: 857
d: [2,2896,2,19,6,2,44,2872,927]
d: [2,2932,2,17,5,3,21,2896,757]
d: [2,2933,2,15,4,4,21,2932,1701]
a: [3,2941,2,3,1,2933]
s: 2941
i: [1,550,0,21,9,0,0,255] d5_xxreal_3
d: [2,729,0,19,8,0,18,550]
a: [3,861,0,17,7,729]
s: 861
i: [1,551,0,21,9,0,0,255] d5_xxreal_3
d: [2,728,0,19,8,0,18,551]
a: [3,865,0,17,7,728]
s: 865
a: [3,2917,2,4,1,2914]
s: 2917
d: [2,3117,3,0,0,5,52,752,1714,747,1727,2917,497]
e: 3117
d: [2,3118,3,8,4,1,40,2917,497]
d: [2,3131,3,6,3,2,21,3118,1727]
d: [2,3132,3,4,2,3,21,3131,747]
d: [2,3133,3,2,1,4,21,3132,1714]
d: [2,3134,3,0,0,5,21,3133,752]
e: 3134
r: 1714
r: 1717
r: 1701
r: 1704
i: [1,560,0,21,9,0,0,255] d5_xxreal_3
d: [2,723,0,19,8,0,18,560]
a: [3,881,0,17,7,723]
s: 881
a: [3,1700,1,2,1,1678]
s: 1700
i: [1,561,0,21,9,0,0,255] d5_xxreal_3
d: [2,722,0,19,8,0,18,561]
a: [3,885,0,17,7,722]
s: 885
i: [1,568,0,21,9,0,0,255] d5_xxreal_3
d: [2,717,0,19,8,0,18,568]
a: [3,893,0,17,7,717]
s: 893
a: [3,1705,1,2,1,1678]
s: 1705
i: [1,569,0,21,9,0,0,255] d5_xxreal_3
d: [2,716,0,19,8,0,18,569]
a: [3,897,0,17,7,716]
s: 897
a: [3,1713,1,2,1,1679]
s: 1713
i: [1,650,0,19,7,0,0,255] d5_xxreal_3
s: 650
i: [1,651,0,19,7,0,0,255] d5_xxreal_3
s: 651
a: [3,1718,1,2,1,1679]
s: 1718
i: [1,652,0,19,7,0,0,255] d5_xxreal_3
s: 652
a: [3,2921,2,4,1,2914]
s: 2921
r: 1700
r: 1705
r: 1713
r: 1718
r: 1727
r: 1730
r: 2921
i: [1,653,0,19,7,0,0,255] d5_xxreal_3
s: 653
i: [1,658,0,19,7,0,0,255] d5_xxreal_3
s: 658
s: 1701
i: [1,659,0,19,7,0,0,255] d5_xxreal_3
s: 659
s: 1704
i: [1,660,0,19,7,0,0,255] d5_xxreal_3
s: 660
i: [1,661,0,19,7,0,0,255] d5_xxreal_3
s: 661
s: 1714
i: [1,668,0,19,6,0,0,255] d5_xxreal_3
s: 668
s: 1717
d: [2,3402,0,17,5,0,21,668,306]
d: [2,3403,0,15,4,0,21,3402,306]
s: 3403
i: [1,669,0,19,6,0,0,255] d5_xxreal_3
s: 669
a: [3,1726,1,2,1,1680]
s: 1726
d: [2,3424,0,17,5,0,21,669,306]
d: [2,3425,0,15,4,0,21,3424,306]
s: 3425
a: [3,1731,1,2,1,1680]
s: 1731
i: [1,633,0,23,9,0,0,255] d5_xxreal_3
d: [2,676,0,19,7,0,18,633]
s: 676
i: [1,632,0,23,9,0,0,255] d5_xxreal_3
d: [2,677,0,19,7,0,18,632]
s: 677
a: [3,2922,2,4,1,2914]
s: 2922
i: [1,597,0,23,9,0,0,255] d5_xxreal_3
d: [2,702,0,19,7,0,18,597]
s: 702
a: [3,2925,2,4,1,2914]
s: 2925
i: [1,596,0,23,9,0,0,255] d5_xxreal_3
d: [2,703,0,19,7,0,18,596]
s: 703
i: [1,552,0,21,9,0,0,255] d5_xxreal_3
a: [3,869,0,19,8,552]
s: 869
d: [2,2891,2,15,6,3,44,2872,2480]
d: [2,2904,2,13,5,4,21,2891,752]
d: [2,2905,2,11,4,5,21,2904,747]
d: [2,2906,2,9,3,6,21,2905,1714]
d: [2,2907,2,7,2,7,21,2906,1727]
a: [3,2911,2,5,1,2907]
s: 2911
d: [2,3470,3,0,0,8,52,1194,1701,742,1189,757,2911,2484]
e: 3470
d: [2,3471,3,0,0,8,52,1194,1701,1199,1189,757,2911,2480]
e: 3471
d: [2,3482,3,13,6,3,44,2911,2480]
d: [2,3485,3,10,5,3,19,3482]
d: [2,3494,3,8,4,4,21,3485,757]
d: [2,3495,3,6,3,5,21,3494,1189]
d: [2,3496,3,4,2,6,21,3495,1199]
d: [2,3497,3,2,1,7,21,3496,1701]
d: [2,3498,3,0,0,8,21,3497,1194]
e: 3498
d: [2,3483,3,13,6,3,44,2911,2484]
d: [2,3484,3,10,5,3,19,3483]
d: [2,3500,3,8,4,4,21,3484,757]
d: [2,3501,3,6,3,5,21,3500,1189]
d: [2,3502,3,4,2,6,21,3501,742]
d: [2,3503,3,2,1,7,21,3502,1701]
d: [2,3504,3,0,0,8,21,3503,1194]
e: 3504
r: 2911
i: [1,553,0,21,9,0,0,255] d5_xxreal_3
a: [3,873,0,19,8,553]
s: 873
d: [2,2886,2,15,6,2,44,2872,923]
d: [2,3186,2,13,5,3,21,2886,752]
d: [2,3187,2,11,4,4,21,3186,747]
d: [2,3188,2,9,3,5,21,3187,1727]
a: [3,3192,2,5,1,3188]
s: 3192
i: [1,666,0,20,6,0,0,255] d5_xxreal_3
s: 666
d: [2,3542,0,18,5,0,21,666,306]
d: [2,3543,0,16,4,0,21,3542,306]
s: 3543
d: [2,3116,3,5,2,2,44,2930,1415]
s: 3116
i: [1,667,0,20,6,0,0,255] d5_xxreal_3
s: 667
d: [2,3506,3,9,3,1,44,3192,499]
d: [2,3522,3,7,2,2,21,3506,787]
d: [2,3523,3,5,1,3,21,3522,757]
a: [3,3527,3,5,1,3523]
s: 3527
d: [2,3567,0,18,5,0,21,667,306]
d: [2,3568,0,16,4,0,21,3567,306]
s: 3568
i: [1,631,0,23,9,0,0,255] d5_xxreal_3
d: [2,678,0,21,8,0,18,631]
s: 678
d: [2,2420,1,23,6,2,44,737,931]
d: [2,3376,1,19,5,3,21,2420,1986]
d: [2,3377,1,17,4,4,21,3376,747]
d: [2,3378,1,13,3,5,21,3377,2031]
a: [3,3382,1,7,1,3378]
s: 3382
d: [2,3615,3,0,0,7,52,757,1701,742,2922,2925,3382,931]
e: 3615
d: [2,3631,2,17,6,2,44,3382,931]
d: [2,3636,2,14,5,2,19,3631]
d: [2,3653,2,12,4,3,21,3636,757]
d: [2,3654,2,8,3,4,21,3653,2925]
d: [2,3655,2,6,2,5,21,3654,742]
d: [2,3656,2,4,1,6,21,3655,1701]
d: [2,3657,2,0,0,7,21,3656,2922]
e: 3657
r: 2031
r: 2034
r: 3382
i: [1,630,0,23,9,0,0,255] d5_xxreal_3
d: [2,679,0,21,8,0,18,630]
s: 679
a: [3,2030,2,4,1,2026]
s: 2030
d: [2,3666,3,0,0,5,52,757,1701,752,1714,2030,497]
e: 3666
d: [2,3667,3,8,4,1,40,2030,497]
d: [2,3680,3,6,3,2,21,3667,1714]
d: [2,3681,3,4,2,3,21,3680,752]
d: [2,3682,3,2,1,4,21,3681,1701]
d: [2,3683,3,0,0,5,21,3682,757]
e: 3683
d: [2,3678,3,12,6,3,44,2030,2484]
d: [2,3690,3,10,5,4,21,3678,1194]
d: [2,3691,3,8,4,5,21,3690,757]
d: [2,3692,3,6,3,6,21,3691,752]
d: [2,3693,3,4,2,7,21,3692,742]
d: [2,3694,3,2,1,8,21,3693,1701]
d: [2,3695,3,0,0,9,21,3694,1714]
e: 3695
r: 1709
r: 1714
r: 2030
i: [1,629,0,23,9,0,0,255] d5_xxreal_3
d: [2,680,0,21,8,0,18,629]
s: 680
i: [1,628,0,23,9,0,0,255] d5_xxreal_3
d: [2,681,0,21,8,0,18,628]
s: 681
s: 1709
i: [1,623,0,23,9,0,0,255] d5_xxreal_3
d: [2,682,0,21,8,0,18,623]
s: 682
s: 1714
i: [1,622,0,23,9,0,0,255] d5_xxreal_3
d: [2,683,0,21,8,0,18,622]
s: 683
i: [1,621,0,23,9,0,0,255] d5_xxreal_3
d: [2,684,0,21,8,0,18,621]
s: 684
d: [2,1810,1,6,3,1,40,449,752]
a: [3,1831,1,2,1,1810]
s: 1831
i: [1,620,0,23,9,0,0,255] d5_xxreal_3
d: [2,685,0,21,8,0,18,620]
s: 685
d: [2,1811,1,6,3,1,40,449,747]
a: [3,1836,1,2,1,1811]
s: 1836
i: [1,619,0,23,9,0,0,255] d5_xxreal_3
d: [2,686,0,21,8,0,18,619]
s: 686
i: [1,618,0,23,9,0,0,255] d5_xxreal_3
d: [2,687,0,21,8,0,18,618]
s: 687
d: [2,1883,2,2,1,1,40,1710,309]
a: [3,1893,2,2,1,1883]
s: 1893
i: [1,617,0,23,9,0,0,255] d5_xxreal_3
d: [2,688,0,21,8,0,18,617]
s: 688
d: [2,1884,2,2,1,1,40,1710,308]
a: [3,1898,2,2,1,1884]
s: 1898
i: [1,616,0,23,9,0,0,255] d5_xxreal_3
d: [2,689,0,21,8,0,18,616]
s: 689
i: [1,613,0,23,9,0,0,255] d5_xxreal_3
d: [2,690,0,21,8,0,18,613]
s: 690
d: [2,1885,2,2,1,1,40,1710,307]
a: [3,1903,2,2,1,1885]
s: 1903
i: [1,612,0,23,9,0,0,255] d5_xxreal_3
d: [2,691,0,21,8,0,18,612]
s: 691
d: [2,3714,2,2,1,1,40,1831,305]
a: [3,3718,2,2,1,3714]
s: 3718
i: [1,611,0,23,9,0,0,255] d5_xxreal_3
d: [2,692,0,21,8,0,18,611]
s: 692
i: [1,610,0,23,9,0,0,255] d5_xxreal_3
d: [2,693,0,21,8,0,18,610]
s: 693
d: [2,3721,2,2,1,1,40,1836,305]
a: [3,3725,2,2,1,3721]
s: 3725
i: [1,609,0,23,9,0,0,255] d5_xxreal_3
d: [2,694,0,21,8,0,18,609]
s: 694
d: [2,3727,3,2,1,1,40,1893,302]
a: [3,3732,3,2,1,3727]
s: 3732
i: [1,608,0,23,9,0,0,255] d5_xxreal_3
d: [2,695,0,21,8,0,18,608]
s: 695
i: [1,607,0,23,9,0,0,255] d5_xxreal_3
d: [2,696,0,21,8,0,18,607]
s: 696
d: [2,3728,3,2,1,1,40,1893,301]
a: [3,3737,3,2,1,3728]
s: 3737
i: [1,606,0,23,9,0,0,255] d5_xxreal_3
d: [2,697,0,21,8,0,18,606]
s: 697
s: 2031
i: [1,601,0,23,9,0,0,255] d5_xxreal_3
d: [2,698,0,21,8,0,18,601]
s: 698
i: [1,600,0,23,9,0,0,255] d5_xxreal_3
d: [2,699,0,21,8,0,18,600]
s: 699
s: 2034
i: [1,599,0,23,9,0,0,255] d5_xxreal_3
d: [2,700,0,21,8,0,18,599]
s: 700
a: [3,2039,2,4,1,2026]
s: 2039
d: [2,3781,3,12,6,3,44,2039,2484]
d: [2,3787,3,10,5,4,21,3781,1199]
d: [2,3788,3,8,4,5,21,3787,757]
d: [2,3789,3,6,3,6,21,3788,752]
d: [2,3790,3,4,2,7,21,3789,742]
d: [2,3791,3,2,1,8,21,3790,1701]
d: [2,3792,3,0,0,9,21,3791,1714]
e: 3792
r: 1709
r: 1714
r: 1726
r: 1731
r: 2031
r: 1701
r: 1704
r: 2034
r: 2039
r: 2872
r: 1717
r: 3192
r: 3527
i: [1,598,0,23,9,0,0,255] d5_xxreal_3
d: [2,701,0,21,8,0,18,598]
s: 701
i: [1,602,0,23,9,0,0,255] d5_xxreal_3
s: 602
s: 1700
i: [1,603,0,23,9,0,0,255] d5_xxreal_3
s: 603
s: 1705
i: [1,626,0,23,9,0,0,255] d5_xxreal_3
s: 626
i: [1,627,0,23,9,0,0,255] d5_xxreal_3
s: 627
s: 1709
i: [1,664,0,23,8,0,0,255] d5_xxreal_3
s: 664
s: 1713
i: [1,665,0,29,6,0,0,255] d5_xxreal_3
s: 665
d: [2,3843,0,27,5,0,21,665,306]
d: [2,3844,0,25,4,0,21,3843,306]
s: 3844
s: 1718
s: 1727
s: 1730
a: [3,2871,1,7,1,2868]
s: 2871
r: 737
d: [2,3888,0,7,1,2,23,737,2871]
a: [3,3928,0,7,1,3888]
s: 3928
s: 2034
d: [2,3937,1,15,6,3,44,3928,2479]
d: [2,3956,1,13,5,4,21,3937,757]
d: [2,3957,1,11,4,5,21,3956,752]
d: [2,3958,1,9,3,6,21,3957,742]
d: [2,3959,1,7,2,7,21,3958,1705]
d: [2,3960,1,5,1,8,21,3959,1718]
a: [3,3964,1,5,1,3960]
s: 3964
d: [2,3991,2,0,0,8,52,1727,1194,1199,747,1189,3964,2484]
e: 3991
d: [2,4001,2,13,6,3,44,3964,2484]
d: [2,4002,2,10,5,3,19,4001]
d: [2,4007,2,8,4,4,21,4002,1189]
d: [2,4008,2,6,3,5,21,4007,747]
d: [2,4009,2,4,2,6,21,4008,1199]
d: [2,4010,2,2,1,7,21,4009,1194]
d: [2,4011,2,0,0,8,21,4010,1727]
e: 4011
r: 1700
r: 1705
r: 2034
r: 3964
r: 2871
s: 1701
s: 737
s: 1704
s: 2030
d: [2,3935,1,15,6,2,44,3928,931]
d: [2,4013,1,13,5,3,21,3935,757]
d: [2,4014,1,11,4,4,21,4013,752]
d: [2,4015,1,9,3,5,21,4014,742]
d: [2,4016,1,7,2,6,21,4015,1718]
a: [3,4020,1,5,1,4016]
s: 4020
a: [3,2035,2,4,1,2026]
s: 2035
s: 2872
d: [2,3940,1,13,3,1,44,3928,499]
d: [2,3966,1,11,2,2,21,3940,747]
d: [2,3967,1,7,1,3,21,3966,1986]
a: [3,3971,1,7,1,3967]
s: 3971
s: 2039
d: [2,4223,3,12,6,2,44,2039,931]
d: [2,4229,3,10,5,3,21,4223,767]
d: [2,4230,3,8,4,4,21,4229,757]
d: [2,4231,3,6,3,5,21,4230,752]
d: [2,4232,3,4,2,6,21,4231,742]
d: [2,4233,3,2,1,7,21,4232,1701]
d: [2,4234,3,0,0,8,21,4233,1718]
e: 4234
r: 1701
r: 1704
r: 1713
r: 1718
r: 1727
r: 1730
r: 2872
r: 1709
r: 2035
r: 2922
r: 4020
s: 1700
s: 1705
s: 1710
s: 1713
s: 1717
s: 1726
s: 1731
a: [3,4019,1,5,1,4016]
s: 4019
d: [2,3930,1,15,6,2,44,3928,849]
d: [2,4342,1,13,5,3,21,3930,757]
d: [2,4343,1,11,4,4,21,4342,752]
d: [2,4344,1,9,3,5,21,4343,742]
a: [3,4348,1,5,1,4344]
s: 4348
s: 2034
s: 2871
r: 737
r: 2871
d: [2,2027,2,12,3,1,40,1986,441]
a: [3,2044,2,4,1,2027]
s: 2044
s: 737
s: 2872
s: 2921
d: [2,2261,1,20,10,1,44,475,919]
d: [2,2286,1,16,8,1,18,2261]
d: [2,2293,1,14,7,1,21,2286,446]
d: [2,2294,1,12,6,1,21,2293,447]
a: [3,2304,1,2,1,2294]
s: 2304
a: [3,2937,2,4,1,2933]
s: 2937
d: [2,2264,1,20,10,1,44,489,919]
d: [2,2283,1,16,8,1,18,2264]
d: [2,2306,1,14,7,1,21,2283,456]
d: [2,2307,1,12,6,1,21,2306,456]
a: [3,2311,1,2,1,2307]
s: 2311
d: [2,2833,1,19,6,3,44,737,2480]
d: [2,2874,1,17,5,4,21,2833,757]
d: [2,2875,1,15,4,5,21,2874,752]
d: [2,2876,1,13,3,6,21,2875,1701]
d: [2,2877,1,11,2,7,21,2876,1714]
a: [3,2881,1,9,1,2877]
s: 2881
d: [2,3771,3,4,1,1,40,2039,305]
a: [3,3785,3,4,1,3771]
s: 3785
d: [2,2176,1,19,6,2,44,737,849]
d: [2,3165,1,17,5,3,21,2176,757]
d: [2,3166,1,15,4,4,21,3165,752]
d: [2,3167,1,13,3,5,21,3166,742]
a: [3,3171,1,9,1,3167]
s: 3171
d: [2,3459,3,12,3,1,40,2925,441]
d: [2,4357,3,8,2,2,21,3459,2917]
a: [3,4361,3,4,1,4357]
s: 4361
a: [3,1991,1,11,1,1983]
s: 1991
d: [2,1417,1,3,1,0,40,436,467]
a: [3,1422,1,3,1,1417]
s: 1422
d: [2,4509,3,4,1,1,40,2044,309]
a: [3,4525,3,4,1,4509]
s: 4525
d: [2,1424,1,3,1,0,40,437,467]
a: [3,1429,1,3,1,1424]
s: 1429
d: [2,4510,3,4,1,1,40,2044,308]
a: [3,4530,3,4,1,4510]
s: 4530
d: [2,1431,1,3,1,0,40,438,467]
a: [3,1436,1,3,1,1431]
s: 1436
d: [2,1849,1,5,2,1,40,461,1205]
d: [2,1861,1,3,1,2,21,1849,837]
a: [3,1865,1,3,1,1861]
s: 1865
d: [2,4511,3,4,1,1,40,2044,307]
a: [3,4535,3,4,1,4511]
s: 4535
a: [3,1870,1,3,1,1851]
s: 1870
d: [2,4616,3,4,1,1,40,2937,305]
a: [3,4630,3,4,1,4616]
s: 4630
d: [2,1908,1,5,2,1,40,466,1407]
a: [3,1912,1,3,1,1908]
s: 1912
d: [2,4875,2,0,0,2,40,1912,1865]
e: 4875
d: [2,4873,2,0,0,2,52,1865,1912]
e: 4873
d: [2,4874,2,0,0,3,52,1205,837,1912,461]
e: 4874
r: 1912
d: [2,2111,1,11,4,1,40,470,1413]
d: [2,2114,1,7,3,2,21,2111,1407]
d: [2,2115,1,5,2,3,21,2114,837]
a: [3,2119,1,3,1,2115]
s: 2119
d: [2,4711,4,4,1,1,40,4361,309]
a: [3,4727,4,4,1,4711]
s: 4727
d: [2,1560,1,4,2,2,44,841,1415]
s: 1560
d: [2,4712,4,4,1,1,40,4361,308]
a: [3,4732,4,4,1,4712]
s: 4732
d: [2,1848,1,7,2,1,40,461,1407]
a: [3,1855,1,4,1,1848]
s: 1855
a: [3,1858,1,3,1,1848]
s: 1858
d: [2,4713,4,4,1,1,40,4361,307]
a: [3,4737,4,4,1,4713]
s: 4737
d: [2,1442,1,5,2,2,44,1205,1415]
s: 1442
d: [2,4793,4,4,1,1,40,4525,302]
a: [3,4808,4,4,1,4793]
s: 4808
d: [2,1599,1,5,2,0,40,303,467]
s: 1599
d: [2,1444,1,6,2,2,44,1413,1415]
s: 1444
d: [2,4794,4,4,1,1,40,4525,301]
a: [3,4813,4,4,1,4794]
s: 4813
d: [2,1653,1,6,2,0,40,296,467]
s: 1653
d: [2,4892,5,4,1,1,40,4727,302]
a: [3,4907,5,4,1,4892]
s: 4907
d: [2,1850,1,6,2,0,40,461,467]
s: 1850
d: [2,1915,1,6,2,1,40,465,1407]
s: 1915
d: [2,4893,5,4,1,1,40,4727,301]
a: [3,4912,5,4,1,4893]
s: 4912
d: [2,1438,1,7,3,1,44,1415,1415]
s: 1438
d: [2,2884,2,15,6,2,44,2872,911]
d: [2,3157,2,13,5,3,21,2884,752]
d: [2,3158,2,11,4,4,21,3157,747]
a: [3,3162,2,5,1,3158]
s: 3162
d: [2,1439,1,7,2,2,44,827,1415]
s: 1439
d: [2,1651,1,7,3,0,40,296,462]
s: 1651
a: [3,3191,2,5,1,3188]
s: 3191
d: [2,1909,1,7,2,0,40,466,467]
s: 1909
d: [2,4187,2,15,6,2,44,3971,931]
d: [2,4206,2,13,5,3,21,4187,757]
d: [2,4207,2,11,4,4,21,4206,752]
d: [2,4208,2,9,3,5,21,4207,742]
d: [2,4209,2,7,2,6,21,4208,1701]
d: [2,4210,2,5,1,7,21,4209,1718]
a: [3,4214,2,5,1,4210]
s: 4214
d: [2,5102,3,0,0,8,52,777,1731,767,787,747,4214,2478]
e: 5102
d: [2,5103,3,9,3,1,44,4214,499]
d: [2,5120,3,4,2,2,21,5103,4019]
d: [2,5121,3,2,1,3,21,5120,787]
d: [2,5122,3,0,0,4,21,5121,747]
e: 5122
d: [2,5104,3,9,3,1,44,4214,499]
d: [2,5124,3,4,2,2,21,5104,4019]
d: [2,5125,3,2,1,3,21,5124,747]
d: [2,5126,3,0,0,4,21,5125,787]
e: 5126
d: [2,5111,3,13,6,3,44,4214,2478]
d: [2,5118,3,10,5,3,19,5111]
d: [2,5128,3,8,4,4,21,5118,747]
d: [2,5129,3,6,3,5,21,5128,787]
d: [2,5130,3,4,2,6,21,5129,767]
d: [2,5131,3,2,1,7,21,5130,1731]
d: [2,5132,3,0,0,8,21,5131,777]
e: 5132
r: 4214
d: [2,1952,1,11,4,0,44,482,304]
d: [2,1965,1,9,3,0,18,1952]
a: [3,1968,1,7,2,1965]
s: 1968
d: [2,2125,1,11,4,1,40,472,1407]
d: [2,2127,1,9,3,2,21,2125,837]
a: [3,2130,1,7,2,2127]
s: 2130
d: [2,4182,2,15,6,2,44,3971,849]
d: [2,4320,2,13,5,3,21,4182,757]
d: [2,4321,2,11,4,4,21,4320,752]
d: [2,4322,2,9,3,5,21,4321,742]
a: [3,4326,2,5,1,4322]
s: 4326
d: [2,2134,1,11,4,0,44,2133,304]
d: [2,2152,1,9,3,0,21,2134,305]
a: [3,2155,1,7,2,2152]
s: 2155
d: [2,4642,2,21,6,2,44,2881,849]
d: [2,4652,2,19,5,3,21,4642,757]
d: [2,4653,2,15,4,4,21,4652,2925]
d: [2,4654,2,13,3,5,21,4653,742]
d: [2,4655,2,9,2,6,21,4654,2921]
d: [2,4656,2,5,1,7,21,4655,2917]
a: [3,4660,2,5,1,4656]
s: 4660
d: [2,1638,1,8,3,0,40,462,303]
s: 1638
d: [2,1878,1,8,3,0,40,483,306]
s: 1878
d: [2,5289,3,13,6,2,44,4660,931]
d: [2,5296,3,11,5,3,21,5289,1189]
d: [2,5297,3,9,4,4,21,5296,747]
d: [2,5298,3,7,3,5,21,5297,1199]
d: [2,5299,3,5,2,6,21,5298,1194]
d: [2,5300,3,3,1,7,21,5299,1731]
a: [3,5304,3,3,1,5300]
s: 5304
d: [2,1879,1,8,3,0,40,483,305]
s: 1879
d: [2,4688,2,21,6,2,44,3171,849]
d: [2,4698,2,19,5,3,21,4688,757]
d: [2,4699,2,15,4,4,21,4698,2925]
d: [2,4700,2,13,3,5,21,4699,742]
d: [2,4701,2,9,2,6,21,4700,2921]
d: [2,4702,2,5,1,7,21,4701,2917]
a: [3,4706,2,5,1,4702]
s: 4706
d: [2,1907,1,8,3,0,40,466,462]
s: 1907
d: [2,1916,1,8,2,0,40,465,467]
s: 1916
a: [3,3526,3,5,1,3523]
s: 3526
d: [2,1997,1,10,4,1,40,1927,464]
a: [3,2000,1,8,3,1997]
s: 2000
r: 1927
d: [2,5369,4,0,0,4,52,841,2930,5304,2000]
e: 5369
r: 2000
d: [2,5032,3,9,3,1,44,3162,499]
d: [2,5048,3,7,2,2,21,5032,841]
d: [2,5049,3,5,1,3,21,5048,757]
a: [3,5053,3,5,1,5049]
s: 5053
s: 1927
d: [2,2146,1,12,5,0,44,483,2133]
d: [2,2147,1,8,3,0,18,2146]
s: 2147
d: [2,5318,4,5,2,2,44,5304,1415]
s: 5318
d: [2,2143,1,12,5,0,44,481,2133]
d: [2,2165,1,10,4,0,21,2143,305]
d: [2,2166,1,8,3,0,21,2165,305]
s: 2166
d: [2,4426,2,7,2,2,44,4348,1415]
s: 4426
d: [2,1625,1,9,3,1,44,304,1415]
s: 1625
d: [2,2053,1,12,6,0,44,479,304]
d: [2,2056,1,10,5,0,21,2053,305]
a: [3,2059,1,4,2,2056]
s: 2059
r: 480
r: 479
r: 478
d: [2,5519,3,0,0,3,52,2030,2039,2059]
e: 5519
d: [2,5520,3,0,0,3,52,2917,2937,2059]
e: 5520
d: [2,5521,2,0,0,3,52,1713,1831,2059]
e: 5521
d: [2,5522,2,0,0,3,52,1726,1836,2059]
e: 5522
d: [2,5523,2,0,0,3,52,1465,1075,2059]
e: 5523
d: [2,5524,3,4,1,2,40,2059,2039]
d: [2,5543,3,0,0,3,21,5524,2030]
e: 5543
d: [2,5525,3,4,1,2,40,2059,2937]
d: [2,5545,3,0,0,3,21,5525,2917]
e: 5545
d: [2,5528,2,2,1,2,40,2059,1831]
d: [2,5547,2,0,0,3,21,5528,1713]
e: 5547
d: [2,5529,2,2,1,2,40,2059,1836]
d: [2,5549,2,0,0,3,21,5529,1726]
e: 5549
d: [2,5536,2,2,1,2,40,2059,1075]
d: [2,5551,2,0,0,3,21,5536,1465]
e: 5551
r: 2059
d: [2,5031,3,7,2,2,44,3162,1415]
s: 5031
s: 478
d: [2,5172,3,7,2,2,44,4326,1415]
s: 5172
s: 479
s: 480
d: [2,5328,3,7,2,2,44,4706,1415]
s: 5328
d: [2,1914,1,9,3,0,40,465,462]
s: 1914
d: [2,5449,3,11,4,2,44,4426,499]
d: [2,5468,3,9,3,3,21,5449,747]
d: [2,5469,3,7,2,4,21,5468,1560]
s: 5469
d: [2,1847,1,10,3,0,40,461,462]
s: 1847
d: [2,1922,1,10,4,0,40,484,309]
s: 1922
d: [2,5450,3,11,4,2,44,4426,499]
d: [2,5470,3,9,3,3,21,5450,747]
d: [2,5471,3,7,2,4,21,5470,1560]
s: 5471
d: [2,5656,1,8,3,2,21,1922,1560]
s: 5656
d: [2,5385,4,7,2,2,44,5053,1415]
s: 5385
d: [2,1923,1,10,4,0,40,484,308]
s: 1923
d: [2,5715,1,8,3,2,21,1923,1560]
s: 5715
d: [2,5553,4,11,4,2,44,5031,499]
d: [2,5572,4,9,3,3,21,5553,757]
d: [2,5573,4,7,2,4,21,5572,1560]
s: 5573
d: [2,1924,1,10,4,0,40,484,307]
s: 1924
d: [2,5554,4,11,4,2,44,5031,499]
d: [2,5574,4,9,3,3,21,5554,757]
d: [2,5575,4,7,2,4,21,5574,1560]
s: 5575
d: [2,5754,1,8,3,2,21,1924,1560]
s: 5754
d: [2,1920,1,14,5,0,40,484,447]
d: [2,1925,1,10,4,0,21,1920,498]
s: 1925
d: [2,5632,4,11,4,2,44,5328,499]
d: [2,5646,4,9,3,3,21,5632,747]
d: [2,5647,4,7,2,4,21,5646,1560]
s: 5647
d: [2,5802,5,7,3,5,44,5647,2155]
d: [2,5803,5,4,2,5,19,5802]
d: [2,5804,5,2,1,6,21,5803,1181]
d: [2,5805,5,0,0,7,21,5804,1836]
e: 5805
r: 1836
r: 2917
r: 5647
r: 5328
r: 4706
r: 4361
d: [2,2177,1,21,6,2,44,737,849]
d: [2,3337,1,17,5,3,21,2177,1986]
d: [2,3338,1,15,4,4,21,3337,747]
a: [3,3342,1,7,1,3338]
s: 3342
a: [3,2918,2,4,1,2914]
s: 2918
s: 1836
d: [2,5789,1,8,3,2,21,1925,1560]
s: 5789
a: [3,4360,3,4,1,4357]
s: 4360
d: [2,5996,4,0,0,4,52,1710,752,747,4360,484]
e: 5996
d: [2,5997,4,6,3,1,40,4360,484]
d: [2,6009,4,4,2,2,21,5997,747]
d: [2,6010,4,2,1,3,21,6009,752]
d: [2,6011,4,0,0,4,21,6010,1710]
e: 6011
r: 1710
r: 1713
r: 1722
r: 1731
r: 2034
r: 2039
r: 4019
r: 2044
r: 4530
r: 4660
s: 1709
s: 1714
s: 1723
s: 1730
d: [2,1930,2,2,1,1,40,1723,309]
a: [3,1940,2,2,1,1930]
s: 1940
s: 4020
d: [2,1931,2,2,1,1,40,1723,308]
a: [3,1945,2,2,1,1931]
s: 1945
d: [2,2419,1,19,6,2,44,737,931]
d: [2,3198,1,17,5,3,21,2419,757]
d: [2,3199,1,15,4,4,21,3198,752]
d: [2,3200,1,13,3,5,21,3199,742]
d: [2,3201,1,11,2,6,21,3200,1701]
a: [3,3205,1,9,1,3201]
s: 3205
d: [2,1921,1,14,5,0,40,484,446]
d: [2,1926,1,10,4,0,21,1921,498]
s: 1926
d: [2,1932,2,2,1,1,40,1723,307]
a: [3,1950,2,2,1,1932]
s: 1950
d: [2,6188,1,8,3,2,21,1926,1560]
s: 6188
d: [2,6128,3,2,1,1,40,1940,302]
a: [3,6133,3,2,1,6128]
s: 6133
d: [2,1977,1,14,6,0,44,484,499]
d: [2,1980,1,10,4,0,18,1977]
s: 1980
d: [2,6228,1,8,3,2,21,1980,1560]
s: 6228
d: [2,6229,4,0,0,5,52,1723,752,4360,6228]
e: 6229
d: [2,6231,4,4,2,3,40,6228,4360]
d: [2,6269,4,2,1,4,21,6231,752]
d: [2,6270,4,0,0,5,21,6269,1723]
e: 6270
r: 1723
r: 1726
s: 1722
s: 1727
d: [2,6129,3,2,1,1,40,1940,301]
a: [3,6138,3,2,1,6129]
s: 6138
s: 3382
d: [2,6283,3,0,0,7,52,757,1705,742,2918,2925,3382,923]
e: 6283
d: [2,6297,2,17,6,2,44,3382,923]
d: [2,6306,2,14,5,2,19,6297]
d: [2,6309,2,12,4,3,21,6306,757]
d: [2,6310,2,8,3,4,21,6309,2925]
d: [2,6311,2,6,2,5,21,6310,742]
d: [2,6312,2,4,1,6,21,6311,1705]
d: [2,6313,2,0,0,7,21,6312,2918]
e: 6313
r: 1700
r: 1705
r: 1714
r: 1717
r: 2881
r: 2918
r: 3171
r: 3342
r: 3928
r: 4360
r: 3191
r: 3526
s: 1701
s: 1704
s: 1713
s: 1718
a: [3,3341,1,7,1,3338]
s: 3341
r: 737
r: 1991
r: 3205
d: [2,6522,0,7,1,2,23,737,3341]
a: [3,6579,0,7,1,6522]
s: 6579
d: [2,4848,4,12,6,2,44,4535,923]
d: [2,6038,4,10,5,3,21,4848,757]
d: [2,6039,4,8,4,4,21,6038,752]
d: [2,6040,4,6,3,5,21,6039,742]
d: [2,6041,4,4,2,6,21,6040,1705]
a: [3,6045,4,2,1,6041]
s: 6045
d: [2,6523,1,7,1,2,23,1991,3341]
a: [3,6584,1,7,1,6523]
s: 6584
d: [2,4798,4,12,6,2,44,4525,923]
d: [2,6056,4,10,5,3,21,4798,757]
d: [2,6057,4,8,4,4,21,6056,752]
d: [2,6058,4,6,3,5,21,6057,742]
d: [2,6059,4,4,2,6,21,6058,1705]
a: [3,6063,4,2,1,6059]
s: 6063
a: [3,2880,1,9,1,2877]
s: 2880
a: [3,3170,1,9,1,3167]
s: 3170
d: [2,4665,4,12,6,2,44,3785,923]
d: [2,6075,4,10,5,3,21,4665,757]
d: [2,6076,4,8,4,4,21,6075,752]
d: [2,6077,4,6,3,5,21,6076,742]
d: [2,6078,4,4,2,6,21,6077,1705]
a: [3,6082,4,2,1,6078]
s: 6082
d: [2,1974,1,14,6,0,44,484,499]
d: [2,1982,1,10,4,0,18,1974]
s: 1982
d: [2,4986,5,12,6,2,44,4813,923]
d: [2,6018,5,10,5,3,21,4986,757]
d: [2,6019,5,8,4,4,21,6018,752]
d: [2,6020,5,6,3,5,21,6019,742]
d: [2,6021,5,4,2,6,21,6020,1705]
a: [3,6025,5,2,1,6021]
s: 6025
d: [2,1995,1,11,4,1,40,1927,469]
s: 1995
d: [2,1996,1,11,4,1,40,1927,460]
s: 1996
d: [2,4963,5,12,6,2,44,4808,923]
d: [2,6028,5,10,5,3,21,4963,757]
d: [2,6029,5,8,4,4,21,6028,752]
d: [2,6030,5,6,3,5,21,6029,742]
d: [2,6031,5,4,2,6,21,6030,1705]
a: [3,6035,5,2,1,6031]
s: 6035
d: [2,2135,1,15,5,0,44,2133,482]
d: [2,2157,1,13,4,0,21,2135,305]
d: [2,2158,1,11,3,0,21,2157,305]
s: 2158
s: 2035
d: [2,2136,1,15,5,0,44,2133,482]
d: [2,2159,1,13,4,0,21,2136,305]
d: [2,2160,1,11,3,0,21,2159,305]
s: 2160
d: [2,2144,1,15,5,0,44,482,2133]
d: [2,2167,1,13,4,0,21,2144,305]
d: [2,2168,1,11,3,0,21,2167,305]
s: 2168
a: [3,2038,2,4,1,2026]
s: 2038
d: [2,2145,1,15,5,0,44,482,2133]
d: [2,2169,1,13,4,0,21,2145,305]
d: [2,2170,1,11,3,0,21,2169,305]
s: 2170
a: [3,2043,2,4,1,2027]
s: 2043
d: [2,3066,1,17,5,0,44,500,304]
d: [2,3096,1,13,4,0,21,3066,481]
a: [3,3099,1,11,3,3096]
s: 3099
d: [2,5400,1,11,4,1,40,1927,469]
s: 5400
s: 2917
d: [2,5401,1,11,4,1,40,1927,460]
s: 5401
s: 4361
d: [2,6881,4,0,0,5,52,1709,752,1722,747,4361,475]
e: 6881
d: [2,6882,4,8,4,1,40,4361,475]
d: [2,6900,4,6,3,2,21,6882,747]
d: [2,6901,4,4,2,3,21,6900,1722]
d: [2,6902,4,2,1,4,21,6901,752]
d: [2,6903,4,0,0,5,21,6902,1709]
e: 6903
d: [2,1416,1,2,1,1,40,436,1205]
d: [2,1531,2,2,1,1,40,1362,299]
d: [2,1782,2,4,2,1,40,1697,447]
d: [2,1787,2,2,1,2,21,1782,757]
d: [2,1809,1,6,3,1,40,449,757]
d: [2,1826,1,4,2,2,21,1809,742]
d: [2,1929,2,4,2,1,40,1723,446]
d: [2,1935,2,2,1,2,21,1929,747]
d: [2,2137,1,11,4,0,44,304,2133]
d: [2,2161,1,9,3,0,21,2137,305]
d: [2,2094,3,8,4,1,40,495,2034]
d: [2,3231,3,6,3,2,21,2094,752]
d: [2,3235,3,4,2,3,21,3231,757]
d: [2,2319,3,12,6,2,44,2034,923]
d: [2,3275,3,10,5,3,21,2319,777]
d: [2,3276,3,8,4,4,21,3275,757]
d: [2,3277,3,6,3,5,21,3276,752]
d: [2,3278,3,4,2,6,21,3277,742]
d: [2,3331,3,6,3,2,21,3118,747]
d: [2,3332,3,4,2,3,21,3331,752]
d: [2,3119,3,8,4,1,40,2917,493]
d: [2,3367,3,6,3,2,21,3119,747]
d: [2,3368,3,4,2,3,21,3367,752]
d: [2,3399,2,4,2,1,40,1714,456]
d: [2,3400,2,2,1,2,21,3399,752]
d: [2,3446,3,8,4,1,40,2922,491]
d: [2,3800,3,6,3,2,21,3446,747]
d: [2,3801,3,4,2,3,21,3800,752]
d: [2,2759,1,23,6,3,44,737,2479]
d: [2,3835,1,19,5,4,21,2759,1986]
d: [2,3836,1,17,4,5,21,3835,747]
d: [2,3447,3,8,4,1,40,2922,487]
d: [2,3803,3,6,3,2,21,3447,747]
d: [2,3841,3,4,2,3,21,3803,752]
d: [2,3885,2,4,2,1,40,1727,456]
d: [2,3886,2,2,1,2,21,3885,747]
d: [2,3951,1,19,6,3,44,3928,2484]
d: [2,3952,1,16,5,3,19,3951]
d: [2,3976,1,12,4,4,21,3952,1986]
d: [2,3977,1,10,3,5,21,3976,747]
d: [2,3978,1,8,2,6,21,3977,1727]
d: [2,4083,1,23,6,3,44,737,2484]
d: [2,4105,1,19,5,4,21,4083,1986]
d: [2,4106,1,17,4,5,21,4105,747]
d: [2,4107,1,15,3,6,21,4106,1727]
d: [2,4146,3,8,4,1,40,2035,491]
d: [2,4238,3,6,3,2,21,4146,752]
d: [2,4239,3,4,2,3,21,4238,757]
d: [2,4225,3,12,6,3,44,2039,2479]
d: [2,4243,3,10,5,4,21,4225,1199]
d: [2,4244,3,8,4,5,21,4243,757]
d: [2,4245,3,6,3,6,21,4244,752]
d: [2,4246,3,4,2,7,21,4245,742]
d: [2,4221,3,12,6,2,44,2039,923]
d: [2,4248,3,10,5,3,21,4221,767]
d: [2,4249,3,8,4,4,21,4248,757]
d: [2,4250,3,6,3,5,21,4249,752]
d: [2,4251,3,4,2,6,21,4250,742]
d: [2,4227,3,12,6,3,44,2039,2484]
d: [2,4282,3,10,5,4,21,4227,1199]
d: [2,4283,3,8,4,5,21,4282,757]
d: [2,4284,3,6,3,6,21,4283,752]
d: [2,4285,3,4,2,7,21,4284,742]
d: [2,4109,3,8,4,1,40,2030,497]
d: [2,4301,3,6,3,2,21,4109,752]
d: [2,4302,3,4,2,3,21,4301,757]
d: [2,4147,3,8,4,1,40,2035,487]
d: [2,4333,3,6,3,2,21,4147,752]
d: [2,4334,3,4,2,3,21,4333,757]
d: [2,4134,2,13,6,3,44,4020,2479]
d: [2,4139,2,10,5,3,19,4134]
d: [2,4351,2,8,4,4,21,4139,787]
d: [2,4352,2,6,3,5,21,4351,747]
d: [2,4353,2,4,2,6,21,4352,767]
d: [2,4354,2,2,1,7,21,4353,777]
d: [2,3933,1,15,6,2,44,3928,923]
d: [2,4366,1,13,5,3,21,3933,757]
d: [2,4367,1,11,4,4,21,4366,752]
d: [2,4368,1,9,3,5,21,4367,742]
d: [2,4370,2,4,2,1,40,1710,447]
d: [2,4375,2,2,1,2,21,4370,752]
d: [2,4371,2,4,2,1,40,1710,446]
d: [2,4377,2,2,1,2,21,4371,752]
d: [2,4548,1,17,3,1,44,737,499]
d: [2,4568,1,13,2,2,21,4548,1986]
d: [2,4569,1,11,1,3,21,4568,747]
d: [2,4588,2,17,6,3,44,2872,2479]
d: [2,4593,2,14,5,3,19,4588]
d: [2,4599,2,12,4,4,21,4593,757]
d: [2,4600,2,8,3,5,21,4599,2925]
d: [2,4601,2,6,2,6,21,4600,742]
d: [2,4602,2,4,1,7,21,4601,1705]
d: [2,5330,3,13,6,2,44,4706,911]
d: [2,5339,3,10,5,2,19,5330]
d: [2,5823,3,8,4,3,21,5339,841]
d: [2,5824,3,6,3,4,21,5823,747]
d: [2,5825,3,4,2,5,21,5824,1547]
d: [2,5826,3,2,1,6,21,5825,1542]
d: [2,4710,4,8,2,1,40,4361,446]
d: [2,5881,4,4,1,2,21,4710,2925]
d: [2,5902,3,8,2,2,21,3459,2921]
d: [2,5914,2,19,6,2,44,3342,849]
d: [2,5925,2,16,5,2,19,5914]
d: [2,5935,2,14,4,3,21,5925,757]
d: [2,5936,2,10,3,4,21,5935,2925]
d: [2,5937,2,8,2,5,21,5936,742]
d: [2,5938,2,4,1,6,21,5937,2921]
d: [2,4506,3,8,4,1,40,2044,475]
d: [2,6014,3,6,3,2,21,4506,752]
d: [2,6015,3,4,2,3,21,6014,757]
d: [2,6016,3,2,1,4,21,6015,1696]
d: [2,4604,3,8,4,1,40,2921,495]
d: [2,6085,3,6,3,2,21,4604,747]
d: [2,6086,3,4,2,3,21,6085,752]
d: [2,5940,3,8,4,1,40,2918,489]
d: [2,6099,3,6,3,2,21,5940,747]
d: [2,6100,3,4,2,3,21,6099,752]
d: [2,5284,3,13,6,2,44,4660,849]
d: [2,5295,3,10,5,2,19,5284]
d: [2,6102,3,8,4,3,21,5295,1189]
d: [2,6103,3,6,3,4,21,6102,747]
d: [2,6104,3,4,2,5,21,6103,1199]
d: [2,6105,3,2,1,6,21,6104,1726]
d: [2,6118,2,4,2,1,40,1723,447]
d: [2,6123,2,2,1,2,21,6118,747]
d: [2,4552,1,23,6,2,44,737,923]
d: [2,6276,1,19,5,3,21,4552,1986]
d: [2,6277,1,17,4,4,21,6276,747]
d: [2,4753,2,23,6,2,44,1991,927]
d: [2,6279,2,21,5,3,21,4753,747]
d: [2,6280,2,17,4,4,21,6279,1986]
d: [2,4590,2,17,6,3,44,2872,2484]
d: [2,4591,2,14,5,3,19,4590]
d: [2,6434,2,12,4,4,21,4591,757]
d: [2,6435,2,8,3,5,21,6434,2925]
d: [2,6436,2,6,2,6,21,6435,742]
d: [2,4546,1,19,6,3,44,737,2484]
d: [2,6438,1,17,5,4,21,4546,757]
d: [2,6439,1,15,4,5,21,6438,752]
d: [2,6440,1,13,3,6,21,6439,742]
d: [2,4544,1,19,6,3,44,737,2479]
d: [2,6505,1,17,5,4,21,4544,757]
d: [2,6506,1,15,4,5,21,6505,752]
d: [2,6507,1,13,3,6,21,6506,742]
d: [2,4537,1,19,6,2,44,737,849]
d: [2,4559,1,17,5,3,21,4537,757]
d: [2,4560,1,15,4,4,21,4559,752]
d: [2,4561,1,13,3,5,21,4560,742]
d: [2,6519,2,4,2,1,40,1701,456]
d: [2,6520,2,2,1,2,21,6519,757]
f: 6903,6520,4561,6507,6311,6440,6436,6280,6277,6269,6123,6105,6100,6086,6016,6010,5938,5902,4701,5881,5826,5804,4656,4602,4569,4377,4375,4368,4354,4334,4302,4285,4251,4246,4239,4232,4107,3978,3888,3886,3841,3836,3377,3801,3657,3400,3368,3338,3332,3278,3235,3049,2161,2027,2021,1935,1826,1787,1680,1679,1531,1448,1473,1530,1416,1201,1185,1183,1177,657,655,649,648,647,707,732,290,289,288,286,285,283,282,281,280,275,276,277,278,279
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_3__l93_xxreal_3
% SZS output start Proof for xxreal_3__l93_xxreal_3
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (~v1_xreal_0(X0) => k4_xxreal_3(X0,k4_xxreal_3(X1,X2)) = k4_xxreal_3(k4_xxreal_3(X0,X1),X2))))) [input l93_xxreal_3]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (~v1_xreal_0(X0) => k4_xxreal_3(X0,k4_xxreal_3(X1,X2)) = k4_xxreal_3(k4_xxreal_3(X0,X1),X2))))) [negated conjecture 1]
9. ! [X0] : ((~v1_xreal_0(X0) & ~v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (v2_xxreal_0(X0) & v1_xxreal_0(X0))) [input cc1_xxreal_3]
12. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
13. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
20. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) => (v2_xxreal_0(X0) & v1_xxreal_0(X0))) [input cc4_xxreal_0]
22. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
26. ! [X0] : ((v1_xxreal_0(X0) & v1_xboole_0(X0)) => (~v3_xxreal_0(X0) & ~v2_xxreal_0(X0) & v1_xxreal_0(X0))) [input cc7_xxreal_0]
27. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
31. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => k4_xxreal_3(X0,X1) = k4_xxreal_3(X1,X0)) [input commutativity_k4_xxreal_3]
32. k1_numbers = k1_xxreal_0 [input d2_xxreal_0]
34. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (~(~(k4_xxreal_3(X0,X1) = X2 <=> k6_numbers = X2) & ~(((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ~(((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ~(~(k4_xxreal_3(X0,X1) = X2 <=> k2_xxreal_0 = X2) & ((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ~(~(k4_xxreal_3(X0,X1) = X2 <=> k1_xxreal_0 = X2) & ((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & ~(v1_xreal_0(X1) & v1_xreal_0(X0))) & ((v1_xreal_0(X1) & v1_xreal_0(X0)) => (k4_xxreal_3(X0,X1) = X2 <=> ? [X3] : (? [X4] : (k3_xcmplx_0(X3,X4) = X2 & X1 = X4 & X0 = X3 & v1_xcmplx_0(X4)) & v1_xcmplx_0(X3)))))))) [input d5_xxreal_3]
43. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v1_xxreal_0(k4_xxreal_3(X0,X1))) [input dt_k4_xxreal_3]
50. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
51. ~v1_xreal_0(k2_xxreal_0) [input fc1_xreal_0]
52. v1_xxreal_0(k1_xxreal_0) [input fc1_xxreal_0]
55. ! [X0,X1] : ((v3_xxreal_0(X1) & v1_xxreal_0(X1) & v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc24_xxreal_3]
57. ! [X0,X1] : ((v3_xxreal_0(X1) & v1_xxreal_0(X1) & v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc25_xxreal_3]
59. ! [X0,X1] : ((v2_xxreal_0(X1) & v1_xxreal_0(X1) & v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc26_xxreal_3]
60. ! [X0,X1] : ((~v3_xxreal_0(X1) & v1_xxreal_0(X1) & ~v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc27_xxreal_3]
61. ! [X0,X1] : ((~v2_xxreal_0(X1) & v1_xxreal_0(X1) & ~v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc28_xxreal_3]
62. ! [X0,X1] : ((~v3_xxreal_0(X1) & v1_xxreal_0(X1) & ~v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1)))) [input fc29_xxreal_3]
63. ~v1_xreal_0(k1_xxreal_0) [input fc2_xreal_0]
64. v1_xxreal_0(k2_xxreal_0) [input fc2_xxreal_0]
65. ! [X0,X1] : ((v1_xxreal_0(X1) & ~v1_xboole_0(X1) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) => (v1_xxreal_0(k4_xxreal_3(X0,X1)) & ~v1_xboole_0(k4_xxreal_3(X0,X1)))) [input fc36_xxreal_3]
66. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0) & v1_xboole_0(X0)) => (v1_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xboole_0(k4_xxreal_3(X0,X1)))) [input fc37_xxreal_3]
69. v3_xxreal_0(k2_xxreal_0) [input fc4_xxreal_0]
73. ! [X0,X1,X2,X3] : ((v1_xcmplx_0(X3) & v1_xcmplx_0(X2) & v1_xreal_0(X1) & v1_xreal_0(X0)) => ((X1 = X3 & X0 = X2) => k4_xxreal_3(X0,X1) = k3_xcmplx_0(X2,X3))) [input ie4_xxreal_3]
96. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
101. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
103. ! [X0] : (v1_xcmplx_0(X0) => k6_numbers = k3_xcmplx_0(X0,k6_numbers)) [input t2_arithm]
108. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
121. ? [X0] : (? [X1] : (? [X2] : ((k4_xxreal_3(X0,k4_xxreal_3(X1,X2)) != k4_xxreal_3(k4_xxreal_3(X0,X1),X2) & ~v1_xreal_0(X0)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
122. ? [X0] : (? [X1] : (? [X2] : (k4_xxreal_3(X0,k4_xxreal_3(X1,X2)) != k4_xxreal_3(k4_xxreal_3(X0,X1),X2) & ~v1_xreal_0(X0) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 121]
126. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
127. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 13]
130. ! [X0] : (k6_numbers = k3_xcmplx_0(X0,k6_numbers) | ~v1_xcmplx_0(X0)) [ennf transformation 103]
131. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
136. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 108]
139. ! [X0] : (! [X1] : (! [X2] : ((((k4_xxreal_3(X0,X1) = X2 <=> k6_numbers = X2) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> k2_xxreal_0 = X2) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> k1_xxreal_0 = X2) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> ? [X3] : (? [X4] : (k3_xcmplx_0(X3,X4) = X2 & X1 = X4 & X0 = X3 & v1_xcmplx_0(X4)) & v1_xcmplx_0(X3))) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0)))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 34]
140. ! [X0] : (! [X1] : (! [X2] : ((((k4_xxreal_3(X0,X1) = X2 <=> k6_numbers = X2) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> k2_xxreal_0 = X2) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> k1_xxreal_0 = X2) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & ((k4_xxreal_3(X0,X1) = X2 <=> ? [X3] : (? [X4] : (k3_xcmplx_0(X3,X4) = X2 & X1 = X4 & X0 = X3 & v1_xcmplx_0(X4)) & v1_xcmplx_0(X3))) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 139]
145. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 27]
149. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | (v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0))) [ennf transformation 20]
150. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0) | v1_xboole_0(X0)) [flattening 149]
155. ! [X0] : ((~v3_xxreal_0(X0) & ~v2_xxreal_0(X0) & v1_xxreal_0(X0)) | (~v1_xxreal_0(X0) | ~v1_xboole_0(X0))) [ennf transformation 26]
156. ! [X0] : ((~v3_xxreal_0(X0) & ~v2_xxreal_0(X0) & v1_xxreal_0(X0)) | ~v1_xxreal_0(X0) | ~v1_xboole_0(X0)) [flattening 155]
157. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | (v1_xreal_0(X0) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 9]
158. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) | v1_xreal_0(X0) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 157]
163. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 22]
164. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 163]
185. ! [X0,X1] : ((v1_xxreal_0(k4_xxreal_3(X0,X1)) & ~v1_xboole_0(k4_xxreal_3(X0,X1))) | (~v1_xxreal_0(X1) | v1_xboole_0(X1) | ~v1_xxreal_0(X0) | v1_xboole_0(X0))) [ennf transformation 65]
186. ! [X0,X1] : ((v1_xxreal_0(k4_xxreal_3(X0,X1)) & ~v1_xboole_0(k4_xxreal_3(X0,X1))) | ~v1_xxreal_0(X1) | v1_xboole_0(X1) | ~v1_xxreal_0(X0) | v1_xboole_0(X0)) [flattening 185]
201. ! [X0,X1] : ((v1_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xboole_0(k4_xxreal_3(X0,X1))) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xboole_0(X0))) [ennf transformation 66]
202. ! [X0,X1] : ((v1_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xboole_0(k4_xxreal_3(X0,X1))) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xboole_0(X0)) [flattening 201]
203. ! [X0,X1] : ((~v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 62]
204. ! [X0,X1] : ((~v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 203]
205. ! [X0,X1] : ((~v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 60]
206. ! [X0,X1] : ((~v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 205]
207. ! [X0,X1] : ((~v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 61]
208. ! [X0,X1] : ((~v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 207]
209. ! [X0,X1] : ((v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 57]
210. ! [X0,X1] : ((v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | ~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 209]
211. ! [X0,X1] : ((v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 55]
212. ! [X0,X1] : ((v3_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | ~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 211]
213. ! [X0,X1] : ((v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | (~v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 59]
214. ! [X0,X1] : ((v2_xxreal_0(k4_xxreal_3(X0,X1)) & v1_xxreal_0(k4_xxreal_3(X0,X1))) | ~v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 213]
215. ! [X0,X1] : (v1_xxreal_0(k4_xxreal_3(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 43]
216. ! [X0,X1] : (v1_xxreal_0(k4_xxreal_3(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 215]
217. ! [X0,X1] : (k4_xxreal_3(X0,X1) = k4_xxreal_3(X1,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 31]
218. ! [X0,X1] : (k4_xxreal_3(X0,X1) = k4_xxreal_3(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 217]
221. ! [X0,X1,X2,X3] : ((k4_xxreal_3(X0,X1) = k3_xcmplx_0(X2,X3) | (X1 != X3 | X0 != X2)) | (~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 73]
222. ! [X0,X1,X2,X3] : (k4_xxreal_3(X0,X1) = k3_xcmplx_0(X2,X3) | X1 != X3 | X0 != X2 | ~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 221]
223. ? [X0] : (? [X1] : (? [X2] : (k4_xxreal_3(X0,k4_xxreal_3(X1,X2)) != k4_xxreal_3(k4_xxreal_3(X0,X1),X2) & ~v1_xreal_0(X0) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (k4_xxreal_3(sK0,k4_xxreal_3(X1,X2)) != k4_xxreal_3(k4_xxreal_3(sK0,X1),X2) & ~v1_xreal_0(sK0) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
224. ? [X1] : (? [X2] : (k4_xxreal_3(sK0,k4_xxreal_3(X1,X2)) != k4_xxreal_3(k4_xxreal_3(sK0,X1),X2) & ~v1_xreal_0(sK0) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (k4_xxreal_3(sK0,k4_xxreal_3(sK1,X2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),X2) & ~v1_xreal_0(sK0) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
225. ? [X2] : (k4_xxreal_3(sK0,k4_xxreal_3(sK1,X2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),X2) & ~v1_xreal_0(sK0) & v1_xxreal_0(X2)) => (k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) & ~v1_xreal_0(sK0) & v1_xxreal_0(sK2)) [choice axiom]
226. ((k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) & ~v1_xreal_0(sK0) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 122,225,224,223]
227. ! [X0] : (! [X1] : (! [X2] : (((((k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2) & (k6_numbers = X2 | k4_xxreal_3(X0,X1) != X2)) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2) & (k2_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2) & (k1_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | ! [X3] : (! [X4] : (k3_xcmplx_0(X3,X4) != X2 | X1 != X4 | X0 != X3 | ~v1_xcmplx_0(X4)) | ~v1_xcmplx_0(X3))) & (? [X3] : (? [X4] : (k3_xcmplx_0(X3,X4) = X2 & X1 = X4 & X0 = X3 & v1_xcmplx_0(X4)) & v1_xcmplx_0(X3)) | k4_xxreal_3(X0,X1) != X2)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 140]
228. ! [X0] : (! [X1] : (! [X2] : (((((k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2) & (k6_numbers = X2 | k4_xxreal_3(X0,X1) != X2)) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2) & (k2_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2) & (k1_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | ! [X3] : (! [X4] : (k3_xcmplx_0(X3,X4) != X2 | X1 != X4 | X0 != X3 | ~v1_xcmplx_0(X4)) | ~v1_xcmplx_0(X3))) & (? [X5] : (? [X6] : (k3_xcmplx_0(X5,X6) = X2 & X1 = X6 & X0 = X5 & v1_xcmplx_0(X6)) & v1_xcmplx_0(X5)) | k4_xxreal_3(X0,X1) != X2)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [rectify 227]
229. ! [X2,X1,X0] : (? [X5] : (? [X6] : (k3_xcmplx_0(X5,X6) = X2 & X1 = X6 & X0 = X5 & v1_xcmplx_0(X6)) & v1_xcmplx_0(X5)) => (? [X6] : (k3_xcmplx_0(sK3(X0,X1,X2),X6) = X2 & X1 = X6 & sK3(X0,X1,X2) = X0 & v1_xcmplx_0(X6)) & v1_xcmplx_0(sK3(X0,X1,X2)))) [choice axiom]
230. ! [X2,X1,X0] : (? [X6] : (k3_xcmplx_0(sK3(X0,X1,X2),X6) = X2 & X1 = X6 & sK3(X0,X1,X2) = X0 & v1_xcmplx_0(X6)) => (k3_xcmplx_0(sK3(X0,X1,X2),sK4(X0,X1,X2)) = X2 & sK4(X0,X1,X2) = X1 & sK3(X0,X1,X2) = X0 & v1_xcmplx_0(sK4(X0,X1,X2)))) [choice axiom]
231. ! [X0] : (! [X1] : (! [X2] : (((((k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2) & (k6_numbers = X2 | k4_xxreal_3(X0,X1) != X2)) | (((v2_xxreal_0(X1) & v3_xxreal_0(X0)) | (v3_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (((v3_xxreal_0(X1) & v3_xxreal_0(X0)) | (v2_xxreal_0(X1) & v2_xxreal_0(X0))) & (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2) & (k2_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v2_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v3_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2) & (k1_xxreal_0 = X2 | k4_xxreal_3(X0,X1) != X2)) | ((~v3_xxreal_0(X1) | ~v3_xxreal_0(X0)) & (~v2_xxreal_0(X1) | ~v2_xxreal_0(X0))) | (v1_xreal_0(X1) & v1_xreal_0(X0))) & (((k4_xxreal_3(X0,X1) = X2 | ! [X3] : (! [X4] : (k3_xcmplx_0(X3,X4) != X2 | X1 != X4 | X0 != X3 | ~v1_xcmplx_0(X4)) | ~v1_xcmplx_0(X3))) & (((k3_xcmplx_0(sK3(X0,X1,X2),sK4(X0,X1,X2)) = X2 & sK4(X0,X1,X2) = X1 & sK3(X0,X1,X2) = X0 & v1_xcmplx_0(sK4(X0,X1,X2))) & v1_xcmplx_0(sK3(X0,X1,X2))) | k4_xxreal_3(X0,X1) != X2)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [skolemisation 228,230,229]
275. v1_xxreal_0(sK0) [cnf transformation 226] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
276. v1_xxreal_0(sK1) [cnf transformation 226] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
277. v1_xxreal_0(sK2) [cnf transformation 226] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
278. ~v1_xreal_0(sK0) [cnf transformation 226] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
279. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) [cnf transformation 226] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
280. ~v1_xreal_0(k1_xxreal_0) [cnf transformation 63] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
281. ~v1_xreal_0(k2_xxreal_0) [cnf transformation 51] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
282. v1_xboole_0(k1_xboole_0) [cnf transformation 50] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
283. v3_xxreal_0(k2_xxreal_0) [cnf transformation 69] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
285. v1_xxreal_0(k2_xxreal_0) [cnf transformation 64] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
286. v1_xxreal_0(k1_xxreal_0) [cnf transformation 52] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
288. k6_numbers = k1_xboole_0 [cnf transformation 96] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
289. k1_numbers = k1_xxreal_0 [cnf transformation 32] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
290. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 101] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
299. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 126] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
300. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 127] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
304. k6_numbers = k3_xcmplx_0(X0,k6_numbers) | ~v1_xcmplx_0(X0) [cnf transformation 130] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
305. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 131] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
310. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 136] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
322. k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2 | ~v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 231] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
324. k4_xxreal_3(X0,X1) = X2 | k1_xxreal_0 != X2 | ~v3_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 231] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
330. k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2 | ~v3_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 231] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
331. k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2 | ~v3_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 231] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
332. k4_xxreal_3(X0,X1) = X2 | k2_xxreal_0 != X2 | ~v2_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 231] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:3,thDist:-3}
397. k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2 | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 231] {a:0,w:24,wCS:9600,sine:255,thAx:0,allAx:3,thDist:-3}
432. k4_xxreal_3(X0,X1) = X2 | k6_numbers != X2 | v2_xxreal_0(X1) | v3_xxreal_0(X1) | v3_xxreal_0(X1) | v2_xxreal_0(X1) | v1_xreal_0(X0) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 231] {a:0,w:24,wCS:9600,sine:255,thAx:0,allAx:3,thDist:-3}
436. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 145] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
441. ~v1_xxreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xboole_0(X0) [cnf transformation 150] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
446. ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) | ~v2_xxreal_0(X0) [cnf transformation 156] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
447. ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) | ~v3_xxreal_0(X0) [cnf transformation 156] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
449. ~v1_xxreal_0(X0) | v1_xreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 158] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
456. ~v2_xxreal_0(X0) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 164] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
475. ~v1_xboole_0(k4_xxreal_3(X0,X1)) | ~v1_xxreal_0(X1) | v1_xboole_0(X1) | ~v1_xxreal_0(X0) | v1_xboole_0(X0) [cnf transformation 186] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
484. v1_xboole_0(k4_xxreal_3(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xboole_0(X0) [cnf transformation 202] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
487. ~v3_xxreal_0(k4_xxreal_3(X0,X1)) | v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 204] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
489. ~v2_xxreal_0(k4_xxreal_3(X0,X1)) | v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 206] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
491. ~v3_xxreal_0(k4_xxreal_3(X0,X1)) | v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 208] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
493. v2_xxreal_0(k4_xxreal_3(X0,X1)) | ~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 210] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
495. v3_xxreal_0(k4_xxreal_3(X0,X1)) | ~v3_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 212] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
497. v2_xxreal_0(k4_xxreal_3(X0,X1)) | ~v2_xxreal_0(X1) | ~v1_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 214] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
498. v1_xxreal_0(k4_xxreal_3(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 216] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
499. k4_xxreal_3(X0,X1) = k4_xxreal_3(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 218] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
501. k4_xxreal_3(X0,X1) = k3_xcmplx_0(X2,X3) | X1 != X3 | X0 != X2 | ~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 222] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:1,thDist:-1}
547. k4_xxreal_3(X0,X1) = k6_numbers | v2_xxreal_0(X1) | v3_xxreal_0(X1) | v3_xxreal_0(X1) | v2_xxreal_0(X1) | v1_xreal_0(X0) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 432] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:3,thDist:-3}
582. k4_xxreal_3(X0,X1) = k6_numbers | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 397] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:3,thDist:-3}
647. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v2_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 332] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
648. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v3_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 331] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
649. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v3_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 330] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
655. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v3_xxreal_0(X1) | ~v3_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(k1_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 324] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
657. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | v1_xreal_0(X0) | ~v1_xxreal_0(k1_xxreal_0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [equality resolution 322] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
672. k3_xcmplx_0(X2,X3) = k4_xxreal_3(X0,X3) | X0 != X2 | ~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X3) | ~v1_xreal_0(X0) [equality resolution 501] {a:0,w:18,wCS:7200,sine:255,thAx:0,allAx:1,thDist:-1}
673. k3_xcmplx_0(X2,X3) = k4_xxreal_3(X2,X3) | ~v1_xcmplx_0(X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) [equality resolution 672] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:1,thDist:-1}
707. k4_xxreal_3(X0,X1) = k6_numbers | v3_xxreal_0(X0) | v2_xxreal_0(X0) | v1_xreal_0(X1) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [duplicate literal removal 582] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
732. k4_xxreal_3(X0,X1) = k6_numbers | v2_xxreal_0(X1) | v3_xxreal_0(X1) | v1_xreal_0(X0) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [duplicate literal removal 547] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:3,thDist:-3}
735. 1 <=> k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) = k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) [avatar definition]
737. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) <- (~1) [avatar component clause 735] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
738. ~1 [avatar split clause 279,735]
740. 2 <=> v1_xreal_0(sK0) [avatar definition]
742. ~v1_xreal_0(sK0) <- (~2) [avatar component clause 740] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
743. ~2 [avatar split clause 278,740]
745. 3 <=> v1_xxreal_0(sK2) [avatar definition]
747. v1_xxreal_0(sK2) <- (3) [avatar component clause 745] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
748. 3 [avatar split clause 277,745]
750. 4 <=> v1_xxreal_0(sK1) [avatar definition]
752. v1_xxreal_0(sK1) <- (4) [avatar component clause 750] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
753. 4 [avatar split clause 276,750]
755. 5 <=> v1_xxreal_0(sK0) [avatar definition]
757. v1_xxreal_0(sK0) <- (5) [avatar component clause 755] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
758. 5 [avatar split clause 275,755]
760. 6 <=> v1_xreal_0(k1_xxreal_0) [avatar definition]
762. ~v1_xreal_0(k1_xxreal_0) <- (~6) [avatar component clause 760] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
763. ~6 [avatar split clause 280,760]
765. 7 <=> v1_xreal_0(k2_xxreal_0) [avatar definition]
767. ~v1_xreal_0(k2_xxreal_0) <- (~7) [avatar component clause 765] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
768. ~7 [avatar split clause 281,765]
770. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
772. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 770] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
773. 8 [avatar split clause 282,770]
775. 9 <=> v3_xxreal_0(k2_xxreal_0) [avatar definition]
777. v3_xxreal_0(k2_xxreal_0) <- (9) [avatar component clause 775] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
778. 9 [avatar split clause 283,775]
785. 11 <=> v1_xxreal_0(k2_xxreal_0) [avatar definition]
787. v1_xxreal_0(k2_xxreal_0) <- (11) [avatar component clause 785] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
788. 11 [avatar split clause 285,785]
790. 12 <=> v1_xxreal_0(k1_xxreal_0) [avatar definition]
792. v1_xxreal_0(k1_xxreal_0) <- (12) [avatar component clause 790] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
793. 12 [avatar split clause 286,790]
800. 14 <=> k6_numbers = k1_xboole_0 [avatar definition]
802. k6_numbers = k1_xboole_0 <- (14) [avatar component clause 800] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
803. 14 [avatar split clause 288,800]
805. 15 <=> k1_numbers = k1_xxreal_0 [avatar definition]
807. k1_numbers = k1_xxreal_0 <- (15) [avatar component clause 805] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
808. 15 [avatar split clause 289,805]
810. 16 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
812. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (16) [avatar component clause 810] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
813. 16 [avatar split clause 290,810]
840. 22 <=> v1_xxreal_0(k6_numbers) [avatar definition]
841. v1_xxreal_0(k6_numbers) <- (22) [avatar component clause 840] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
848. 24 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k6_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | v3_xxreal_0(X1) | v2_xxreal_0(X1)) [avatar definition]
849. k4_xxreal_3(X0,X1) = k6_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | v3_xxreal_0(X1) | v2_xxreal_0(X1) <- (24) [avatar component clause 848] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
850. ~22 | 24 [avatar split clause 732,848,840]
910. 35 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k6_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | v2_xxreal_0(X0) | v3_xxreal_0(X0)) [avatar definition]
911. k4_xxreal_3(X0,X1) = k6_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | v2_xxreal_0(X0) | v3_xxreal_0(X0) <- (35) [avatar component clause 910] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
912. ~22 | 35 [avatar split clause 707,910,840]
922. 38 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v2_xxreal_0(X1)) [avatar definition]
923. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v2_xxreal_0(X1) <- (38) [avatar component clause 922] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
924. ~11 | 38 [avatar split clause 647,922,785]
926. 39 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v2_xxreal_0(X0) | ~v3_xxreal_0(X1)) [avatar definition]
927. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X1) | ~v2_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (39) [avatar component clause 926] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
928. ~11 | 39 [avatar split clause 648,926,785]
930. 40 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v3_xxreal_0(X1)) [avatar definition]
931. k4_xxreal_3(X0,X1) = k2_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (40) [avatar component clause 930] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
932. ~11 | 40 [avatar split clause 649,930,785]
938. 42 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v3_xxreal_0(X1)) [avatar definition]
939. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (42) [avatar component clause 938] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:3,thDist:-3}
940. ~12 | 42 [avatar split clause 655,938,790]
946. 44 <=> ! [X1,X0] : (k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v2_xxreal_0(X1)) [avatar definition]
947. k4_xxreal_3(X0,X1) = k1_xxreal_0 | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v2_xxreal_0(X1) <- (44) [avatar component clause 946] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:3,thDist:-3}
948. ~12 | 44 [avatar split clause 657,946,790]
1177. v1_xboole_0(k6_numbers) <- (8, 14) [backward demodulation 772,802] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
1179. 91 <=> v1_xboole_0(k6_numbers) [avatar definition]
1181. v1_xboole_0(k6_numbers) <- (91) [avatar component clause 1179] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1182. 91 | ~8 | ~14 [avatar split clause 1177,800,770,1179]
1183. v1_xxreal_0(k1_numbers) <- (12, 15) [backward demodulation 792,807] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
1185. ~v1_xreal_0(k1_numbers) <- (~6, 15) [backward demodulation 762,807] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
1187. 92 <=> v1_xxreal_0(k1_numbers) [avatar definition]
1189. v1_xxreal_0(k1_numbers) <- (92) [avatar component clause 1187] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1190. 92 | ~12 | ~15 [avatar split clause 1183,805,790,1187]
1197. 94 <=> v1_xreal_0(k1_numbers) [avatar definition]
1199. ~v1_xreal_0(k1_numbers) <- (~94) [avatar component clause 1197] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1200. ~94 | 6 | ~15 [avatar split clause 1185,805,760,1197]
1201. m1_subset_1(k6_numbers,k4_ordinal1) <- (14, 16) [forward demodulation 812,802] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
1203. 95 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
1205. m1_subset_1(k6_numbers,k4_ordinal1) <- (95) [avatar component clause 1203] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1206. 95 | ~14 | ~16 [avatar split clause 1201,810,800,1203]
1360. 118 <=> v7_ordinal1(k6_numbers) [avatar definition]
1362. v7_ordinal1(k6_numbers) <- (118) [avatar component clause 1360] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1415. k6_numbers = X0 | ~v1_xboole_0(X0) <- (14) [forward demodulation 310,802] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
1416. v7_ordinal1(k6_numbers) <- (95) [resolution 436,1205] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
1418. 118 | ~95 [avatar split clause 1416,1203,1360]
1448. ~v1_xxreal_0(k6_numbers) | ~v2_xxreal_0(k6_numbers) <- (91) [resolution 446,1181] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:2,thDist:-2}
1473. ~v1_xxreal_0(k6_numbers) | ~v3_xxreal_0(k6_numbers) <- (91) [resolution 447,1181] {a:1,w:4,wCS:1600,sine:2,thAx:0,allAx:2,thDist:-2}
1530. v1_xxreal_0(k6_numbers) <- (118) [resolution 1362,300] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1531. v1_xreal_0(k6_numbers) <- (118) [resolution 1362,299] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
1538. 22 | ~118 [avatar split clause 1530,1360,840]
1540. 143 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1542. ~v3_xxreal_0(k6_numbers) <- (~143) [avatar component clause 1540] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1543. ~143 | ~22 | ~91 [avatar split clause 1473,1179,840,1540]
1545. 144 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1547. ~v2_xxreal_0(k6_numbers) <- (~144) [avatar component clause 1545] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
1548. ~144 | ~22 | ~91 [avatar split clause 1448,1179,840,1545]
1556. 146 <=> v1_xreal_0(k6_numbers) [avatar definition]
1559. 146 | ~118 [avatar split clause 1531,1360,1556]
1560. ~v1_xboole_0(X0) | v1_xxreal_0(X0) <- (14, 22) [superposition 841,1415] {a:1,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:4,thDist:-4}
1679. v3_xxreal_0(sK1) | v2_xxreal_0(sK1) | v1_xboole_0(sK1) <- (4) [resolution 441,752] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1680. v3_xxreal_0(sK2) | v2_xxreal_0(sK2) | v1_xboole_0(sK2) <- (3) [resolution 441,747] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1695. 163 <=> v1_xboole_0(sK0) [avatar definition]
1696. ~v1_xboole_0(sK0) <- (~163) [avatar component clause 1695] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1697. v1_xboole_0(sK0) <- (163) [avatar component clause 1695] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1699. 164 <=> v2_xxreal_0(sK0) [avatar definition]
1701. v2_xxreal_0(sK0) <- (164) [avatar component clause 1699] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1703. 165 <=> v3_xxreal_0(sK0) [avatar definition]
1705. v3_xxreal_0(sK0) <- (165) [avatar component clause 1703] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1708. 166 <=> v1_xboole_0(sK1) [avatar definition]
1709. ~v1_xboole_0(sK1) <- (~166) [avatar component clause 1708] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1710. v1_xboole_0(sK1) <- (166) [avatar component clause 1708] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1712. 167 <=> v2_xxreal_0(sK1) [avatar definition]
1714. v2_xxreal_0(sK1) <- (167) [avatar component clause 1712] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1716. 168 <=> v3_xxreal_0(sK1) [avatar definition]
1719. 166 | 167 | 168 | ~4 [avatar split clause 1679,750,1716,1712,1708]
1721. 169 <=> v1_xboole_0(sK2) [avatar definition]
1722. ~v1_xboole_0(sK2) <- (~169) [avatar component clause 1721] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1723. v1_xboole_0(sK2) <- (169) [avatar component clause 1721] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1725. 170 <=> v2_xxreal_0(sK2) [avatar definition]
1726. ~v2_xxreal_0(sK2) <- (~170) [avatar component clause 1725] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1727. v2_xxreal_0(sK2) <- (170) [avatar component clause 1725] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1729. 171 <=> v3_xxreal_0(sK2) [avatar definition]
1732. 169 | 170 | 171 | ~3 [avatar split clause 1680,745,1729,1725,1721]
1782. ~v1_xxreal_0(sK0) | ~v3_xxreal_0(sK0) <- (163) [resolution 1697,447] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1787. ~v3_xxreal_0(sK0) <- (5, 163) [subsumption resolution 1782,757] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1788. ~165 | ~5 | ~163 [avatar split clause 1787,1695,755,1703]
1809. v1_xreal_0(sK0) | v3_xxreal_0(sK0) | v2_xxreal_0(sK0) <- (5) [resolution 449,757] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1826. v3_xxreal_0(sK0) | v2_xxreal_0(sK0) <- (~2, 5) [subsumption resolution 1809,742] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1827. 164 | 165 | 2 | ~5 [avatar split clause 1826,755,740,1703,1699]
1834. 187 <=> v1_xreal_0(sK2) [avatar definition]
1929. ~v1_xxreal_0(sK2) | ~v2_xxreal_0(sK2) <- (169) [resolution 1723,446] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1935. ~v2_xxreal_0(sK2) <- (3, 169) [subsumption resolution 1929,747] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1936. ~170 | ~3 | ~169 [avatar split clause 1935,1721,745,1725]
1977. v1_xboole_0(k4_xxreal_3(X1,X0)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) [superposition 484,499] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
1980. v1_xboole_0(k4_xxreal_3(X1,X0)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | ~v1_xboole_0(X0) [duplicate literal removal 1977] {a:1,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
1985. 202 <=> v1_xxreal_0(k4_xxreal_3(sK0,sK1)) [avatar definition]
1986. v1_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (202) [avatar component clause 1985] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1987. ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (~202) [avatar component clause 1985] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1989. 203 <=> k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) = k4_xxreal_3(sK2,k4_xxreal_3(sK0,sK1)) [avatar definition]
1991. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(sK2,k4_xxreal_3(sK0,sK1)) <- (~203) [avatar component clause 1989] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
2018. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~202) [resolution 1987,498] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2020. ~v1_xxreal_0(sK0) <- (4, ~202) [subsumption resolution 2018,752] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2021. $false <- (4, 5, ~202) [subsumption resolution 2020,757] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
2022. ~4 | ~5 | 202 [avatar contradiction clause 2021]
2027. v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | v1_xboole_0(k4_xxreal_3(sK0,sK1)) <- (202) [resolution 1986,441] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
2029. 205 <=> v2_xxreal_0(k4_xxreal_3(sK0,sK1)) [avatar definition]
2030. ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (~205) [avatar component clause 2029] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2033. 206 <=> v3_xxreal_0(k4_xxreal_3(sK0,sK1)) [avatar definition]
2034. ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (~206) [avatar component clause 2033] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2035. v3_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (206) [avatar component clause 2033] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2037. 207 <=> v1_xreal_0(k4_xxreal_3(sK0,sK1)) [avatar definition]
2039. v1_xreal_0(k4_xxreal_3(sK0,sK1)) <- (207) [avatar component clause 2037] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2042. 208 <=> v1_xboole_0(k4_xxreal_3(sK0,sK1)) [avatar definition]
2044. v1_xboole_0(k4_xxreal_3(sK0,sK1)) <- (208) [avatar component clause 2042] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
2045. 208 | 205 | 206 | ~202 [avatar split clause 2027,1985,2033,2029,2042]
2094. ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (~206) [resolution 495,2034] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
2132. k3_xcmplx_0(X2,X3) = k4_xxreal_3(X2,X3) | ~v1_xcmplx_0(X2) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) [subsumption resolution 673,305] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:1,thDist:-1}
2133. k3_xcmplx_0(X2,X3) = k4_xxreal_3(X2,X3) | ~v1_xreal_0(X3) | ~v1_xreal_0(X2) [subsumption resolution 2132,305] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
2137. k6_numbers = k4_xxreal_3(X0,k6_numbers) | ~v1_xcmplx_0(X0) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X0) [superposition 304,2133] {a:1,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
2154. 212 <=> ! [X0] : (k6_numbers = k4_xxreal_3(X0,k6_numbers) | ~v1_xreal_0(X0)) [avatar definition]
2155. k6_numbers = k4_xxreal_3(X0,k6_numbers) | ~v1_xreal_0(X0) <- (212) [avatar component clause 2154] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
2161. k6_numbers = k4_xxreal_3(X0,k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(X0) [subsumption resolution 2137,305] {a:1,w:9,wCS:3600,sine:255,thAx:0,allAx:2,thDist:-2}
2162. ~146 | 212 [avatar split clause 2161,2154,1556]
2177. k6_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (~1, 24) [superposition 737,849] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2319. ~v3_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (38, ~206) [superposition 2034,923] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
2420. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (~1, 40) [superposition 737,931] {a:1,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2479. k4_xxreal_3(X0,X1) = k1_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v3_xxreal_0(X0) | ~v3_xxreal_0(X1) <- (15, 42) [forward demodulation 939,807] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
2484. k4_xxreal_3(X0,X1) = k1_numbers | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X1) | v1_xreal_0(X0) | ~v2_xxreal_0(X0) | ~v2_xxreal_0(X1) <- (15, 44) [forward demodulation 947,807] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
2759. k1_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (~1, 15, 42) [superposition 737,2479] {a:1,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2870. 227 <=> k1_numbers = k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) [avatar definition]
2871. k1_numbers = k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) <- (227) [avatar component clause 2870] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2872. k1_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) <- (~227) [avatar component clause 2870] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2879. 228 <=> k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) = k4_xxreal_3(k1_numbers,sK2) [avatar definition]
2881. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k1_numbers,sK2) <- (~228) [avatar component clause 2879] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
2916. 230 <=> v2_xxreal_0(k4_xxreal_3(sK1,sK2)) [avatar definition]
2917. ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~230) [avatar component clause 2916] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
2918. v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (230) [avatar component clause 2916] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
2920. 231 <=> v3_xxreal_0(k4_xxreal_3(sK1,sK2)) [avatar definition]
2921. ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~231) [avatar component clause 2920] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
2922. v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (231) [avatar component clause 2920] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
2924. 232 <=> v1_xxreal_0(k4_xxreal_3(sK1,sK2)) [avatar definition]
2925. v1_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (232) [avatar component clause 2924] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
2926. ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~232) [avatar component clause 2924] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3036. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~232) [resolution 2926,498] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3048. ~v1_xxreal_0(sK1) <- (3, ~232) [subsumption resolution 3036,747] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3049. $false <- (3, 4, ~232) [subsumption resolution 3048,752] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3050. ~3 | ~4 | 232 [avatar contradiction clause 3049]
3118. ~v2_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (~230) [resolution 2917,497] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3119. ~v3_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (~230) [resolution 2917,493] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3169. 239 <=> k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) = k4_xxreal_3(k6_numbers,sK2) [avatar definition]
3171. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k6_numbers,sK2) <- (~239) [avatar component clause 3169] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3231. ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (4, ~206) [subsumption resolution 2094,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
3235. ~v3_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (4, 5, ~206) [subsumption resolution 3231,757] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
3236. ~164 | ~168 | ~4 | ~5 | 206 [avatar split clause 3235,2033,755,750,1716,1699]
3275. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (9, 38, ~206) [subsumption resolution 2319,777] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
3276. ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (5, 9, 38, ~206) [subsumption resolution 3275,757] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
3277. v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (4, 5, 9, 38, ~206) [subsumption resolution 3276,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
3278. ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (~2, 4, 5, 9, 38, ~206) [subsumption resolution 3277,742] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
3279. ~167 | ~165 | 2 | ~4 | ~5 | ~9 | ~38 | 206 [avatar split clause 3278,2033,922,775,755,750,740,1703,1712]
3331. ~v2_xxreal_0(sK2) | ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (3, ~230) [subsumption resolution 3118,747] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3332. ~v2_xxreal_0(sK2) | ~v2_xxreal_0(sK1) <- (3, 4, ~230) [subsumption resolution 3331,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3333. ~167 | ~170 | ~3 | ~4 | 230 [avatar split clause 3332,2916,750,745,1725,1712]
3337. k6_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (~1, 24, 202) [subsumption resolution 2177,1986] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3338. k6_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (~1, 3, 24, 202) [subsumption resolution 3337,747] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3340. 242 <=> k6_numbers = k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) [avatar definition]
3342. k6_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) <- (~242) [avatar component clause 3340] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3343. 170 | 171 | 207 | ~242 | 1 | ~3 | ~24 | ~202 [avatar split clause 3338,1985,848,745,735,3340,2037,1729,1725]
3367. ~v3_xxreal_0(sK2) | ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (3, ~230) [subsumption resolution 3119,747] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3368. ~v3_xxreal_0(sK2) | ~v3_xxreal_0(sK1) <- (3, 4, ~230) [subsumption resolution 3367,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3369. ~168 | ~171 | ~3 | ~4 | 230 [avatar split clause 3368,2916,750,745,1729,1716]
3376. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (~1, 40, 202) [subsumption resolution 2420,1986] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3377. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (~1, 3, 40, 202) [subsumption resolution 3376,747] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3380. 243 <=> k2_xxreal_0 = k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) [avatar definition]
3382. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) <- (~243) [avatar component clause 3380] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3399. ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (167) [resolution 1714,456] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3400. ~v3_xxreal_0(sK1) <- (4, 167) [subsumption resolution 3399,752] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3401. ~168 | ~4 | ~167 [avatar split clause 3400,1712,750,1716]
3446. v2_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (231) [resolution 2922,491] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3447. v3_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (231) [resolution 2922,487] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3459. v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xboole_0(k4_xxreal_3(sK1,sK2)) <- (232) [resolution 2925,441] {a:3,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:11,thDist:-11}
3631. k2_xxreal_0 != k2_xxreal_0 | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (40, ~243) [superposition 3382,931] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3636. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (40, ~243) [trivial inequality removal 3631] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3653. ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 40, ~243) [subsumption resolution 3636,757] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3654. v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 40, 232, ~243) [subsumption resolution 3653,2925] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3655. ~v2_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 40, 232, ~243) [subsumption resolution 3654,742] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3656. ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 40, 164, 232, ~243) [subsumption resolution 3655,1701] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3657. $false <- (~2, 5, 40, 164, 231, 232, ~243) [subsumption resolution 3656,2922] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
3658. 2 | ~5 | ~40 | ~164 | ~231 | ~232 | 243 [avatar contradiction clause 3657]
3800. v2_xxreal_0(sK2) | v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (3, 231) [subsumption resolution 3446,747] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3801. v2_xxreal_0(sK2) | v2_xxreal_0(sK1) <- (3, 4, 231) [subsumption resolution 3800,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3802. 167 | 170 | ~3 | ~4 | ~231 [avatar split clause 3801,2920,750,745,1725,1712]
3803. v3_xxreal_0(sK2) | v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (3, 231) [subsumption resolution 3447,747] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3831. ~171 | ~205 | 207 | ~243 | 1 | ~3 | ~40 | ~202 [avatar split clause 3377,1985,930,745,735,3380,2037,2029,1729]
3835. k1_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (~1, 15, 42, 202) [subsumption resolution 2759,1986] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3836. k1_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(sK2) <- (~1, 3, 15, 42, 202) [subsumption resolution 3835,747] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3837. ~171 | ~206 | 207 | ~227 | 1 | ~3 | ~15 | ~42 | ~202 [avatar split clause 3836,1985,938,805,745,735,2870,2037,2033,1729]
3841. v3_xxreal_0(sK2) | v3_xxreal_0(sK1) <- (3, 4, 231) [subsumption resolution 3803,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
3842. 168 | 171 | ~3 | ~4 | ~231 [avatar split clause 3841,2920,750,745,1729,1716]
3885. ~v3_xxreal_0(sK2) | ~v1_xxreal_0(sK2) <- (170) [resolution 1727,456] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3886. ~v3_xxreal_0(sK2) <- (3, 170) [subsumption resolution 3885,747] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
3887. ~171 | ~3 | ~170 [avatar split clause 3886,1725,745,1729]
3888. k1_numbers != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) <- (~1, 227) [backward demodulation 737,2871] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
3926. 250 <=> k1_numbers = k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) [avatar definition]
3928. k1_numbers != k4_xxreal_3(k4_xxreal_3(sK0,sK1),sK2) <- (~250) [avatar component clause 3926] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
3929. ~250 | 1 | ~227 [avatar split clause 3888,2870,735,3926]
3933. k1_numbers != k4_xxreal_3(k2_xxreal_0,sK2) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (38, ~250) [superposition 3928,923] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3951. k1_numbers != k1_numbers | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (15, 44, ~250) [superposition 3928,2484] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3952. ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (15, 44, ~250) [trivial inequality removal 3951] {a:1,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3976. ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (15, 44, 202, ~250) [subsumption resolution 3952,1986] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3977. v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (3, 15, 44, 202, ~250) [subsumption resolution 3976,747] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3978. v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (3, 15, 44, 170, 202, ~250) [subsumption resolution 3977,1727] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
3979. ~205 | 207 | ~3 | ~15 | ~44 | ~170 | ~202 | 250 [avatar split clause 3978,3926,1985,1725,946,805,745,2037,2029]
4018. 253 <=> k1_numbers = k4_xxreal_3(k2_xxreal_0,sK2) [avatar definition]
4020. k1_numbers != k4_xxreal_3(k2_xxreal_0,sK2) <- (~253) [avatar component clause 4018] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4083. k1_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (~1, 15, 44) [superposition 737,2484] {a:1,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4105. k1_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (~1, 15, 44, 202) [subsumption resolution 4083,1986] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4106. k1_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (~1, 3, 15, 44, 202) [subsumption resolution 4105,747] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4107. k1_numbers != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (~1, 3, 15, 44, 170, 202) [subsumption resolution 4106,1727] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4108. ~205 | 207 | ~227 | 1 | ~3 | ~15 | ~44 | ~170 | ~202 [avatar split clause 4107,1985,1725,946,805,745,735,2870,2037,2029]
4109. ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (~205) [resolution 2030,497] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4134. k1_numbers != k1_numbers | ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK2) | v1_xreal_0(k2_xxreal_0) | ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (15, 42, ~253) [superposition 4020,2479] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4139. ~v1_xxreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK2) | v1_xreal_0(k2_xxreal_0) | ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (15, 42, ~253) [trivial inequality removal 4134] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4146. v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) | v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (206) [resolution 2035,491] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4147. v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) | v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (206) [resolution 2035,487] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4221. v1_xreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (38, 207) [superposition 2039,923] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4223. v1_xreal_0(k2_xxreal_0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (40, 207) [superposition 2039,931] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4225. v1_xreal_0(k1_numbers) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (15, 42, 207) [superposition 2039,2479] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4227. v1_xreal_0(k1_numbers) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (15, 44, 207) [superposition 2039,2484] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4229. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (~7, 40, 207) [subsumption resolution 4223,767] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4230. ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (5, ~7, 40, 207) [subsumption resolution 4229,757] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4231. v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (4, 5, ~7, 40, 207) [subsumption resolution 4230,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4232. ~v2_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (~2, 4, 5, ~7, 40, 207) [subsumption resolution 4231,742] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4236. ~168 | ~164 | 2 | ~4 | ~5 | 7 | ~40 | ~207 [avatar split clause 4232,2037,930,765,755,750,740,1699,1716]
4238. v2_xxreal_0(sK1) | v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (4, 206) [subsumption resolution 4146,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4239. v2_xxreal_0(sK1) | v2_xxreal_0(sK0) <- (4, 5, 206) [subsumption resolution 4238,757] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4240. 164 | 167 | ~4 | ~5 | ~206 [avatar split clause 4239,2033,755,750,1712,1699]
4243. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (15, 42, ~94, 207) [subsumption resolution 4225,1199] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4244. ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (5, 15, 42, ~94, 207) [subsumption resolution 4243,757] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4245. v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (4, 5, 15, 42, ~94, 207) [subsumption resolution 4244,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4246. ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (~2, 4, 5, 15, 42, ~94, 207) [subsumption resolution 4245,742] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4247. ~168 | ~165 | 2 | ~4 | ~5 | ~15 | ~42 | 94 | ~207 [avatar split clause 4246,2037,1197,938,805,755,750,740,1703,1716]
4248. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (~7, 38, 207) [subsumption resolution 4221,767] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4249. ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (5, ~7, 38, 207) [subsumption resolution 4248,757] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4250. v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (4, 5, ~7, 38, 207) [subsumption resolution 4249,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4251. ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (~2, 4, 5, ~7, 38, 207) [subsumption resolution 4250,742] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4252. ~167 | ~165 | 2 | ~4 | ~5 | 7 | ~38 | ~207 [avatar split clause 4251,2037,922,765,755,750,740,1703,1712]
4282. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (15, 44, ~94, 207) [subsumption resolution 4227,1199] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4283. ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (5, 15, 44, ~94, 207) [subsumption resolution 4282,757] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4284. v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (4, 5, 15, 44, ~94, 207) [subsumption resolution 4283,752] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4285. ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (~2, 4, 5, 15, 44, ~94, 207) [subsumption resolution 4284,742] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
4286. ~167 | ~164 | 2 | ~4 | ~5 | ~15 | ~44 | 94 | ~207 [avatar split clause 4285,2037,1197,946,805,755,750,740,1699,1712]
4301. ~v2_xxreal_0(sK1) | ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (4, ~205) [subsumption resolution 4109,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4302. ~v2_xxreal_0(sK1) | ~v2_xxreal_0(sK0) <- (4, 5, ~205) [subsumption resolution 4301,757] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4303. ~164 | ~167 | ~4 | ~5 | 205 [avatar split clause 4302,2029,755,750,1712,1699]
4333. v3_xxreal_0(sK1) | v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (4, 206) [subsumption resolution 4147,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4334. v3_xxreal_0(sK1) | v3_xxreal_0(sK0) <- (4, 5, 206) [subsumption resolution 4333,757] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4335. 165 | 168 | ~4 | ~5 | ~206 [avatar split clause 4334,2033,755,750,1716,1703]
4351. ~v1_xxreal_0(sK2) | v1_xreal_0(k2_xxreal_0) | ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (11, 15, 42, ~253) [subsumption resolution 4139,787] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4352. v1_xreal_0(k2_xxreal_0) | ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (3, 11, 15, 42, ~253) [subsumption resolution 4351,747] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4353. ~v3_xxreal_0(k2_xxreal_0) | ~v3_xxreal_0(sK2) <- (3, ~7, 11, 15, 42, ~253) [subsumption resolution 4352,767] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4354. ~v3_xxreal_0(sK2) <- (3, ~7, 9, 11, 15, 42, ~253) [subsumption resolution 4353,777] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4355. ~171 | ~3 | 7 | ~9 | ~11 | ~15 | ~42 | 253 [avatar split clause 4354,4018,938,805,785,775,765,745,1729]
4359. 258 <=> v1_xboole_0(k4_xxreal_3(sK1,sK2)) [avatar definition]
4360. ~v1_xboole_0(k4_xxreal_3(sK1,sK2)) <- (~258) [avatar component clause 4359] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
4361. v1_xboole_0(k4_xxreal_3(sK1,sK2)) <- (258) [avatar component clause 4359] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
4366. k1_numbers != k4_xxreal_3(k2_xxreal_0,sK2) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (5, 38, ~250) [subsumption resolution 3933,757] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4367. k1_numbers != k4_xxreal_3(k2_xxreal_0,sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (4, 5, 38, ~250) [subsumption resolution 4366,752] {a:1,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4368. k1_numbers != k4_xxreal_3(k2_xxreal_0,sK2) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (~2, 4, 5, 38, ~250) [subsumption resolution 4367,742] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4369. ~167 | ~165 | ~253 | 2 | ~4 | ~5 | ~38 | 250 [avatar split clause 4368,3926,922,755,750,740,4018,1703,1712]
4370. ~v1_xxreal_0(sK1) | ~v3_xxreal_0(sK1) <- (166) [resolution 1710,447] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
4371. ~v1_xxreal_0(sK1) | ~v2_xxreal_0(sK1) <- (166) [resolution 1710,446] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
4375. ~v3_xxreal_0(sK1) <- (4, 166) [subsumption resolution 4370,752] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
4376. ~168 | ~4 | ~166 [avatar split clause 4375,1708,750,1716]
4377. ~v2_xxreal_0(sK1) <- (4, 166) [subsumption resolution 4371,752] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
4378. ~167 | ~4 | ~166 [avatar split clause 4377,1708,750,1712]
4506. ~v1_xxreal_0(sK1) | v1_xboole_0(sK1) | ~v1_xxreal_0(sK0) | v1_xboole_0(sK0) <- (208) [resolution 2044,475] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
4537. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k6_numbers,sK2) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | v3_xxreal_0(sK1) | v2_xxreal_0(sK1) <- (~1, 24) [superposition 737,849] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4544. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k1_numbers,sK2) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (~1, 15, 42) [superposition 737,2479] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4546. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k1_numbers,sK2) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (~1, 15, 44) [superposition 737,2484] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4548. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(sK2,k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) <- (~1) [superposition 737,499] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
4552. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (~1, 38) [superposition 737,923] {a:1,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4559. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k6_numbers,sK2) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | v3_xxreal_0(sK1) | v2_xxreal_0(sK1) <- (~1, 5, 24) [subsumption resolution 4537,757] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4560. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k6_numbers,sK2) | v1_xreal_0(sK0) | v3_xxreal_0(sK1) | v2_xxreal_0(sK1) <- (~1, 4, 5, 24) [subsumption resolution 4559,752] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4561. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k6_numbers,sK2) | v3_xxreal_0(sK1) | v2_xxreal_0(sK1) <- (~1, ~2, 4, 5, 24) [subsumption resolution 4560,742] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
4568. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(sK2,k4_xxreal_3(sK0,sK1)) | ~v1_xxreal_0(sK2) <- (~1, 202) [subsumption resolution 4548,1986] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
4569. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(sK2,k4_xxreal_3(sK0,sK1)) <- (~1, 3, 202) [subsumption resolution 4568,747] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
4570. ~203 | 1 | ~3 | ~202 [avatar split clause 4569,1985,745,735,1989]
4588. k1_numbers != k1_numbers | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (15, 42, ~227) [superposition 2872,2479] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4590. k1_numbers != k1_numbers | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (15, 44, ~227) [superposition 2872,2484] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4591. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (15, 44, ~227) [trivial inequality removal 4590] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4593. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (15, 42, ~227) [trivial inequality removal 4588] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4599. ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 15, 42, ~227) [subsumption resolution 4593,757] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4600. v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 15, 42, ~227, 232) [subsumption resolution 4599,2925] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4601. ~v3_xxreal_0(sK0) | ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 15, 42, ~227, 232) [subsumption resolution 4600,742] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4602. ~v3_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 15, 42, 165, ~227, 232) [subsumption resolution 4601,1705] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4603. ~231 | 2 | ~5 | ~15 | ~42 | ~165 | 227 | ~232 [avatar split clause 4602,2924,2870,1703,938,805,755,740,2920]
4604. ~v3_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (~231) [resolution 2921,495] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
4642. k6_numbers != k4_xxreal_3(k1_numbers,sK2) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (24, ~228) [superposition 2881,849] {a:2,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4652. k6_numbers != k4_xxreal_3(k1_numbers,sK2) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 24, ~228) [subsumption resolution 4642,757] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4653. k6_numbers != k4_xxreal_3(k1_numbers,sK2) | v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 24, ~228, 232) [subsumption resolution 4652,2925] {a:2,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4654. k6_numbers != k4_xxreal_3(k1_numbers,sK2) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 24, ~228, 232) [subsumption resolution 4653,742] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4655. k6_numbers != k4_xxreal_3(k1_numbers,sK2) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 24, ~228, ~231, 232) [subsumption resolution 4654,2921] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4656. k6_numbers != k4_xxreal_3(k1_numbers,sK2) <- (~2, 5, 24, ~228, ~230, ~231, 232) [subsumption resolution 4655,2917] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4658. 263 <=> k6_numbers = k4_xxreal_3(k1_numbers,sK2) [avatar definition]
4660. k6_numbers != k4_xxreal_3(k1_numbers,sK2) <- (~263) [avatar component clause 4658] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4661. ~263 | 2 | ~5 | ~24 | 228 | 230 | 231 | ~232 [avatar split clause 4656,2924,2920,2916,2879,848,755,740,4658]
4688. k6_numbers != k4_xxreal_3(k6_numbers,sK2) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (24, ~239) [superposition 3171,849] {a:2,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4698. k6_numbers != k4_xxreal_3(k6_numbers,sK2) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 24, ~239) [subsumption resolution 4688,757] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4699. k6_numbers != k4_xxreal_3(k6_numbers,sK2) | v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 24, 232, ~239) [subsumption resolution 4698,2925] {a:2,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4700. k6_numbers != k4_xxreal_3(k6_numbers,sK2) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 24, 232, ~239) [subsumption resolution 4699,742] {a:2,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4701. k6_numbers != k4_xxreal_3(k6_numbers,sK2) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 24, ~231, 232, ~239) [subsumption resolution 4700,2921] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4704. 264 <=> k6_numbers = k4_xxreal_3(k6_numbers,sK2) [avatar definition]
4706. k6_numbers != k4_xxreal_3(k6_numbers,sK2) <- (~264) [avatar component clause 4704] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
4710. ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (258) [resolution 4361,446] {a:4,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
4753. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (39, ~203) [superposition 1991,927] {a:2,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
5284. k6_numbers != k6_numbers | ~v1_xxreal_0(k1_numbers) | ~v1_xxreal_0(sK2) | v1_xreal_0(k1_numbers) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (24, ~263) [superposition 4660,849] {a:3,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
5295. ~v1_xxreal_0(k1_numbers) | ~v1_xxreal_0(sK2) | v1_xreal_0(k1_numbers) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (24, ~263) [trivial inequality removal 5284] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
5328. k4_xxreal_3(X0,sK2) != X0 | ~v1_xboole_0(X0) <- (14, ~264) [superposition 4706,1415] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
5330. k6_numbers != k6_numbers | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK2) | v2_xxreal_0(k6_numbers) | v3_xxreal_0(k6_numbers) <- (35, ~264) [superposition 4706,911] {a:3,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
5339. ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(sK2) | v1_xreal_0(sK2) | v2_xxreal_0(k6_numbers) | v3_xxreal_0(k6_numbers) <- (35, ~264) [trivial inequality removal 5330] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
5632. k4_xxreal_3(sK2,X0) != X0 | ~v1_xboole_0(X0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X0) <- (14, ~264) [superposition 5328,499] {a:4,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
5646. k4_xxreal_3(sK2,X0) != X0 | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) <- (3, 14, ~264) [subsumption resolution 5632,747] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
5647. k4_xxreal_3(sK2,X0) != X0 | ~v1_xboole_0(X0) <- (3, 14, 22, ~264) [subsumption resolution 5646,1560] {a:4,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
5802. k6_numbers != k6_numbers | ~v1_xboole_0(k6_numbers) | ~v1_xreal_0(sK2) <- (3, 14, 22, 212, ~264) [superposition 5647,2155] {a:5,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
5803. ~v1_xboole_0(k6_numbers) | ~v1_xreal_0(sK2) <- (3, 14, 22, 212, ~264) [trivial inequality removal 5802] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
5804. ~v1_xreal_0(sK2) <- (3, 14, 22, 91, 212, ~264) [subsumption resolution 5803,1181] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
5807. ~187 | ~3 | ~14 | ~22 | ~91 | ~212 | 264 [avatar split clause 5804,4704,2154,1179,840,800,745,1834]
5823. ~v1_xxreal_0(sK2) | v1_xreal_0(sK2) | v2_xxreal_0(k6_numbers) | v3_xxreal_0(k6_numbers) <- (22, 35, ~264) [subsumption resolution 5339,841] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
5824. v1_xreal_0(sK2) | v2_xxreal_0(k6_numbers) | v3_xxreal_0(k6_numbers) <- (3, 22, 35, ~264) [subsumption resolution 5823,747] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
5825. v1_xreal_0(sK2) | v3_xxreal_0(k6_numbers) <- (3, 22, 35, ~144, ~264) [subsumption resolution 5824,1547] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
5826. v1_xreal_0(sK2) <- (3, 22, 35, ~143, ~144, ~264) [subsumption resolution 5825,1542] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
5827. 187 | ~3 | ~22 | ~35 | 143 | 144 | 264 [avatar split clause 5826,4704,1545,1540,910,840,745,1834]
5881. ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (232, 258) [subsumption resolution 4710,2925] {a:4,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
5882. ~230 | ~232 | ~258 [avatar split clause 5881,4359,2924,2916]
5892. 230 | ~264 | 2 | ~5 | ~24 | 231 | ~232 | 239 [avatar split clause 4701,3169,2924,2920,848,755,740,4704,2916]
5902. v2_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xboole_0(k4_xxreal_3(sK1,sK2)) <- (~231, 232) [subsumption resolution 3459,2921] {a:3,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:11,thDist:-11}
5903. 258 | 230 | 231 | ~232 [avatar split clause 5902,2924,2920,2916,4359]
5914. k6_numbers != k6_numbers | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (24, ~242) [superposition 3342,849] {a:2,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
5925. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (24, ~242) [trivial inequality removal 5914] {a:2,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
5935. ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 24, ~242) [subsumption resolution 5925,757] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
5936. v1_xreal_0(sK0) | v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 24, 232, ~242) [subsumption resolution 5935,2925] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
5937. v3_xxreal_0(k4_xxreal_3(sK1,sK2)) | v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 24, 232, ~242) [subsumption resolution 5936,742] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
5938. v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 24, ~231, 232, ~242) [subsumption resolution 5937,2921] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
5939. 230 | 2 | ~5 | ~24 | 231 | ~232 | 242 [avatar split clause 5938,3340,2924,2920,848,755,740,2916]
5940. v3_xxreal_0(sK2) | ~v1_xxreal_0(sK2) | v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (230) [resolution 2918,489] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
5997. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) | ~v1_xboole_0(sK1) <- (~258) [resolution 4360,484] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
6009. ~v1_xxreal_0(sK1) | ~v1_xboole_0(sK1) <- (3, ~258) [subsumption resolution 5997,747] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
6010. ~v1_xboole_0(sK1) <- (3, 4, ~258) [subsumption resolution 6009,752] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
6013. ~166 | ~3 | ~4 | 258 [avatar split clause 6010,4359,750,745,1708]
6014. v1_xboole_0(sK1) | ~v1_xxreal_0(sK0) | v1_xboole_0(sK0) <- (4, 208) [subsumption resolution 4506,752] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
6015. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (4, 5, 208) [subsumption resolution 6014,757] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
6016. v1_xboole_0(sK1) <- (4, 5, ~163, 208) [subsumption resolution 6015,1696] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
6017. 166 | ~4 | ~5 | 163 | ~208 [avatar split clause 6016,2042,1695,755,750,1708]
6085. ~v3_xxreal_0(sK2) | ~v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (3, ~231) [subsumption resolution 4604,747] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
6086. ~v3_xxreal_0(sK2) | ~v2_xxreal_0(sK1) <- (3, 4, ~231) [subsumption resolution 6085,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
6087. ~167 | ~171 | ~3 | ~4 | 231 [avatar split clause 6086,2920,750,745,1729,1712]
6099. v3_xxreal_0(sK2) | v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (3, 230) [subsumption resolution 5940,747] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
6100. v3_xxreal_0(sK2) | v2_xxreal_0(sK1) <- (3, 4, 230) [subsumption resolution 6099,752] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
6101. 167 | 171 | ~3 | ~4 | ~230 [avatar split clause 6100,2916,750,745,1729,1712]
6102. ~v1_xxreal_0(sK2) | v1_xreal_0(k1_numbers) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (24, 92, ~263) [subsumption resolution 5295,1189] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
6103. v1_xreal_0(k1_numbers) | v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (3, 24, 92, ~263) [subsumption resolution 6102,747] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
6104. v3_xxreal_0(sK2) | v2_xxreal_0(sK2) <- (3, 24, 92, ~94, ~263) [subsumption resolution 6103,1199] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
6105. v3_xxreal_0(sK2) <- (3, 24, 92, ~94, ~170, ~263) [subsumption resolution 6104,1726] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:13,thDist:-13}
6106. 171 | ~3 | ~24 | ~92 | 94 | 170 | 263 [avatar split clause 6105,4658,1725,1197,1187,848,745,1729]
6118. ~v1_xxreal_0(sK2) | ~v3_xxreal_0(sK2) <- (169) [resolution 1723,447] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
6123. ~v3_xxreal_0(sK2) <- (3, 169) [subsumption resolution 6118,747] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
6124. ~171 | ~3 | ~169 [avatar split clause 6123,1721,745,1729]
6228. v1_xboole_0(k4_xxreal_3(X1,X0)) | ~v1_xxreal_0(X1) | ~v1_xboole_0(X0) <- (14, 22) [subsumption resolution 1980,1560] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
6231. ~v1_xxreal_0(sK1) | ~v1_xboole_0(sK2) <- (14, 22, ~258) [resolution 6228,4360] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
6269. ~v1_xboole_0(sK2) <- (4, 14, 22, ~258) [subsumption resolution 6231,752] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
6274. ~169 | ~4 | ~14 | ~22 | 258 [avatar split clause 6269,4359,840,800,750,1721]
6276. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(sK2) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (~1, 38, 202) [subsumption resolution 4552,1986] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
6277. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) <- (~1, 3, 38, 202) [subsumption resolution 6276,747] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
6278. ~170 | ~206 | 207 | ~243 | 1 | ~3 | ~38 | ~202 [avatar split clause 6277,1985,922,745,735,3380,2037,2033,1725]
6279. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | ~v1_xxreal_0(k4_xxreal_3(sK0,sK1)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (3, 39, ~203) [subsumption resolution 4753,747] {a:2,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
6280. k2_xxreal_0 != k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) | v1_xreal_0(k4_xxreal_3(sK0,sK1)) | ~v2_xxreal_0(sK2) | ~v3_xxreal_0(k4_xxreal_3(sK0,sK1)) <- (3, 39, 202, ~203) [subsumption resolution 6279,1986] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
6281. ~206 | ~170 | 207 | ~243 | ~3 | ~39 | ~202 | 203 [avatar split clause 6280,1989,1985,926,745,3380,2037,1725,2033]
6297. k2_xxreal_0 != k2_xxreal_0 | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (38, ~243) [superposition 3382,923] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6306. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (38, ~243) [trivial inequality removal 6297] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6309. ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 38, ~243) [subsumption resolution 6306,757] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6310. v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 38, 232, ~243) [subsumption resolution 6309,2925] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6311. ~v3_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 38, 232, ~243) [subsumption resolution 6310,742] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6434. ~v1_xxreal_0(k4_xxreal_3(sK1,sK2)) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 15, 44, ~227) [subsumption resolution 4591,757] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6435. v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (5, 15, 44, ~227, 232) [subsumption resolution 6434,2925] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6436. ~v2_xxreal_0(sK0) | ~v2_xxreal_0(k4_xxreal_3(sK1,sK2)) <- (~2, 5, 15, 44, ~227, 232) [subsumption resolution 6435,742] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
6437. ~230 | ~164 | 2 | ~5 | ~15 | ~44 | 227 | ~232 [avatar split clause 6436,2924,2870,946,805,755,740,1699,2916]
6438. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k1_numbers,sK2) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (~1, 5, 15, 44) [subsumption resolution 4546,757] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
6439. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k1_numbers,sK2) | v1_xreal_0(sK0) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (~1, 4, 5, 15, 44) [subsumption resolution 6438,752] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
6440. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k1_numbers,sK2) | ~v2_xxreal_0(sK0) | ~v2_xxreal_0(sK1) <- (~1, ~2, 4, 5, 15, 44) [subsumption resolution 6439,742] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
6441. ~167 | ~164 | ~228 | 1 | 2 | ~4 | ~5 | ~15 | ~44 [avatar split clause 6440,946,805,755,750,740,735,2879,1699,1712]
6442. ~230 | ~165 | 2 | ~5 | ~38 | ~232 | 243 [avatar split clause 6311,3380,2924,922,755,740,1703,2916]
6505. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k1_numbers,sK2) | ~v1_xxreal_0(sK1) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (~1, 5, 15, 42) [subsumption resolution 4544,757] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
6506. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k1_numbers,sK2) | v1_xreal_0(sK0) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (~1, 4, 5, 15, 42) [subsumption resolution 6505,752] {a:1,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
6507. k4_xxreal_3(sK0,k4_xxreal_3(sK1,sK2)) != k4_xxreal_3(k1_numbers,sK2) | ~v3_xxreal_0(sK0) | ~v3_xxreal_0(sK1) <- (~1, ~2, 4, 5, 15, 42) [subsumption resolution 6506,742] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
6508. ~168 | ~165 | ~228 | 1 | 2 | ~4 | ~5 | ~15 | ~42 [avatar split clause 6507,938,805,755,750,740,735,2879,1703,1716]
6509. 167 | 168 | ~239 | 1 | 2 | ~4 | ~5 | ~24 [avatar split clause 4561,848,755,750,740,735,3169,1716,1712]
6519. ~v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (164) [resolution 1701,456] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
6520. ~v3_xxreal_0(sK0) <- (5, 164) [subsumption resolution 6519,757] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
6521. ~165 | ~5 | ~164 [avatar split clause 6520,1699,755,1703]
6882. ~v1_xxreal_0(sK2) | v1_xboole_0(sK2) | ~v1_xxreal_0(sK1) | v1_xboole_0(sK1) <- (258) [resolution 4361,475] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
6900. v1_xboole_0(sK2) | ~v1_xxreal_0(sK1) | v1_xboole_0(sK1) <- (3, 258) [subsumption resolution 6882,747] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
6901. ~v1_xxreal_0(sK1) | v1_xboole_0(sK1) <- (3, ~169, 258) [subsumption resolution 6900,1722] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
6902. v1_xboole_0(sK1) <- (3, 4, ~169, 258) [subsumption resolution 6901,752] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
6903. $false <- (3, 4, ~166, ~169, 258) [subsumption resolution 6902,1709] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
6904. ~3 | ~4 | 166 | 169 | ~258 [avatar contradiction clause 6903]
6905. $false [avatar sat refutation 738,743,748,753,758,763,768,773,778,788,793,803,808,813,850,912,924,928,932,940,948,1182,1190,1200,1206,1418,1538,1543,1548,1559,1719,1732,1788,1827,1936,2022,2045,2162,3050,3236,3279,3333,3343,3369,3401,3658,3802,3831,3837,3842,3887,3929,3979,4108,4236,4240,4247,4252,4286,4303,4335,4355,4369,4376,4378,4570,4603,4661,5807,5827,5882,5892,5903,5939,6013,6017,6087,6101,6106,6124,6274,6278,6281,6437,6441,6442,6508,6509,6521,6904] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:8010,thDist:-8010}
% SZS output end Proof for xxreal_3__l93_xxreal_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 109
% >>> Preprocessing
% Introduced skolems: 26
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 271
% Generated clauses: 5422
% Active clauses: 402
% Passive clauses: 861
% Final active clauses: 299
% Final passive clauses: 337
% Main loop iterations started: 504
% 
% >>> Simplifying Inferences
% Duplicate literals: 897
% Trivial inequalities: 102
% Fw subsumption resolutions: 1190
% Fw demodulations: 15
% Bw demodulations: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 217
% Equational tautologies: 1
% Forward subsumptions: 1860
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 418
% Unit resulting resolution: 25
% Forward superposition: 1443
% Backward superposition: 530
% Self superposition: 4
% Equality factoring: 1
% Equality resolution: 2
% 
% >>> AVATAR
% Split clauses: 331
% Split components: 845
% Sat splitting refutations: 54
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 710
% SAT solver unit clauses: 68
% SAT solver binary clauses: 175
% 
% Memory used [KB]: 8827
% Time elapsed: 0.031 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% minimizing solver time: 0.001 s
% sat proof minimization: 0.001 s
% splitting model update: 0.001 s
% congruence closure: 0.004 s ( own 0.003 s ) 
% backward subsumption index maintenance: 0.001 s
% forward subsumption: 0.007 s
% forward demodulation: 0.001 s
% unit resulting resolution: 0.001 s
% superposition: 0.005 s
% other: 0.031 s ( own 0.009 s ) 

