i: [1,64,0,4,2,1,0,0] t40_gate_1
a: [3,87,0,2,1,64]
s: 87
a: [3,92,0,2,1,64]
s: 92
i: [1,52,0,7,2,1,0,0] t40_gate_1
a: [3,158,0,2,1,52]
s: 158
i: [1,51,0,7,2,1,0,0] t40_gate_1
a: [3,167,0,2,1,51]
s: 167
i: [1,50,0,7,2,1,0,0] t40_gate_1
a: [3,176,0,2,1,50]
s: 176
i: [1,60,0,6,2,1,0,0] t40_gate_1
a: [3,97,0,4,1,60]
s: 97
i: [1,57,0,6,2,1,0,0] t40_gate_1
a: [3,117,0,4,1,57]
s: 117
i: [1,56,0,6,2,1,0,0] t40_gate_1
a: [3,126,0,4,1,56]
s: 126
i: [1,55,0,6,2,1,0,0] t40_gate_1
a: [3,135,0,4,1,55]
s: 135
i: [1,54,0,6,2,1,0,0] t40_gate_1
a: [3,144,0,4,1,54]
s: 144
i: [1,53,0,7,2,1,0,0] t40_gate_1
a: [3,149,0,5,1,53]
s: 149
i: [1,58,0,9,2,1,0,0] t40_gate_1
a: [3,108,0,7,1,58]
s: 108
i: [1,65,0,2,1,0,0,2] fc1_xboole_0
a: [3,181,0,2,1,65]
s: 181
i: [1,83,0,2,1,0,0,1] rc2_xboole_0
a: [3,186,0,2,1,83]
s: 186
i: [1,84,0,2,1,0,0,1] rc1_xboole_0
a: [3,191,0,2,1,84]
s: 191
i: [1,66,0,5,2,0,0,2] fc2_gate_1
s: 66
i: [1,67,0,5,2,0,0,2] fc1_gate_1
s: 67
i: [1,75,0,6,2,0,0,1] fc5_gate_1
s: 75
i: [1,74,0,7,1,0,0,2] commutativity_k3_gate_1
s: 74
i: [1,76,0,8,3,0,0,1] fc6_gate_1
s: 76
a: [3,103,0,2,1,58]
s: 103
d: [2,197,1,4,2,1,40,76,144]
a: [3,207,1,2,1,197]
s: 207
i: [1,77,0,9,3,0,0,1] fc35_gate_1
s: 77
r: 158
a: [3,157,0,2,1,52]
s: 157
a: [3,154,0,5,1,52]
s: 154
r: 167
a: [3,166,0,2,1,51]
s: 166
d: [2,198,1,4,2,1,40,76,135]
a: [3,216,1,2,1,198]
s: 216
a: [3,163,0,5,1,51]
s: 163
r: 87
r: 92
r: 97
r: 149
r: 154
r: 157
r: 163
r: 166
a: [3,88,0,2,1,64]
s: 88
a: [3,91,0,2,1,64]
s: 91
s: 158
s: 167
a: [3,98,0,4,1,60]
s: 98
d: [2,253,1,0,0,2,52,103,98,75]
e: 253
d: [2,254,1,2,1,1,40,98,75]
d: [2,256,1,0,0,2,21,254,103]
e: 256
r: 88
r: 91
r: 98
r: 158
r: 167
s: 87
s: 92
s: 157
s: 166
s: 97
a: [3,219,1,2,1,198]
s: 219
s: 149
s: 154
d: [2,199,1,4,2,1,40,76,126]
a: [3,229,1,2,1,199]
s: 229
s: 163
d: [2,200,1,4,2,1,40,76,117]
a: [3,238,1,2,1,200]
s: 238
i: [1,78,0,9,3,0,0,1] fc33_gate_1
s: 78
i: [1,79,0,9,3,0,0,1] fc34_gate_1
s: 79
d: [2,269,2,0,0,3,52,216,176,163,79]
e: 269
d: [2,272,1,4,2,1,40,79,163]
d: [2,274,1,2,1,2,21,272,176]
d: [2,275,1,0,0,3,21,274,216]
e: 275
r: 163
r: 166
r: 229
r: 117
r: 154
r: 157
r: 238
a: [3,113,0,2,1,57]
s: 113
s: 158
s: 167
a: [3,116,0,4,1,57]
s: 116
a: [3,153,0,5,1,52]
s: 153
s: 229
i: [1,80,0,9,3,0,0,1] fc38_gate_1
s: 80
a: [3,233,1,2,1,200]
s: 233
i: [1,81,0,9,3,0,0,1] fc36_gate_1
s: 81
i: [1,82,0,9,3,0,0,1] fc37_gate_1
s: 82
a: [3,237,1,2,1,200]
s: 237
i: [1,69,0,10,2,0,0,1] d18_gate_1
s: 69
d: [2,297,1,4,2,1,44,108,69]
d: [2,298,1,2,1,2,21,297,181]
a: [3,302,1,2,1,298]
s: 302
i: [1,70,0,10,2,0,0,1] d18_gate_1
s: 70
a: [3,139,0,2,1,54]
s: 139
d: [2,194,1,6,2,0,44,75,74]
s: 194
i: [1,71,0,10,2,0,0,1] d18_gate_1
s: 71
a: [3,130,0,2,1,55]
s: 130
i: [1,72,0,10,2,0,0,1] d18_gate_1
s: 72
a: [3,121,0,2,1,56]
s: 121
d: [2,195,1,6,2,0,44,75,74]
s: 195
i: [1,73,0,10,2,0,0,1] d18_gate_1
s: 73
d: [2,317,1,4,2,1,44,108,73]
d: [2,318,1,2,1,2,21,317,181]
d: [2,319,1,0,0,3,21,318,113]
e: 319
r: 87
r: 92
r: 97
r: 103
r: 108
s: 88
s: 91
a: [3,104,0,2,1,58]
s: 104
s: 98
d: [2,321,2,0,0,2,52,91,98,194]
e: 321
d: [2,322,2,2,1,1,40,98,194]
d: [2,325,2,0,0,2,21,322,91]
e: 325
r: 88
r: 91
r: 98
r: 104
r: 113
r: 116
r: 153
r: 158
r: 167
r: 176
r: 237
s: 87
s: 92
s: 103
a: [3,112,0,2,1,57]
s: 112
s: 157
s: 166
a: [3,224,1,2,1,199]
s: 224
s: 97
s: 117
s: 238
s: 154
s: 163
a: [3,175,0,2,1,50]
s: 175
a: [3,172,0,5,1,50]
s: 172
d: [2,339,2,0,0,3,52,207,302,172,79]
e: 339
d: [2,340,1,4,2,1,40,172,79]
d: [2,344,1,2,1,2,21,340,302]
d: [2,345,1,0,0,3,21,344,207]
e: 345
r: 207
i: [1,61,0,4,2,1,0,0] t40_gate_1
i: [1,59,0,6,2,1,0,0] t40_gate_1
d: [2,201,1,4,2,1,40,76,97]
d: [2,240,1,2,1,2,21,201,92]
d: [2,250,1,4,2,1,40,163,77]
d: [2,251,1,2,1,2,21,250,176]
d: [2,271,1,4,2,1,40,79,154]
d: [2,260,1,4,2,1,40,154,77]
d: [2,270,1,4,2,1,40,79,149]
d: [2,259,1,4,2,1,40,149,77]
d: [2,304,1,4,2,1,44,108,70]
d: [2,305,1,2,1,2,21,304,181]
d: [2,311,1,4,2,1,44,108,71]
d: [2,312,1,2,1,2,21,311,181]
d: [2,314,1,4,2,1,44,108,72]
d: [2,315,1,2,1,2,21,314,181]
d: [2,328,1,4,2,1,40,117,76]
d: [2,329,1,2,1,2,21,328,233]
d: [2,334,1,4,2,1,40,163,79]
d: [2,337,1,2,1,2,21,334,216]
d: [2,342,1,4,2,1,40,172,77]
d: [2,348,1,2,1,2,21,342,302]
f: 348,344,337,329,325,319,315,312,305,298,259,270,260,271,251,240,199,198,197,65,50,51,52,53,54,55,56,57,58,59,60,61,64
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t40_gate_1
% SZS output start Proof for gate_1__t40_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ~(~(~(v1_xboole_0(X12) & ~v1_xboole_0(X18)) & ~(v1_xboole_0(X18) & ~v1_xboole_0(X12))) & ~(v1_xboole_0(X18) & ~v1_xboole_0(k3_gate_1(X12,X17))) & ~(v1_xboole_0(k3_gate_1(X12,X17)) & ~v1_xboole_0(X18)) & ~(v1_xboole_0(k18_gate_1(X0,X13,X14,X15,X16)) & ~v1_xboole_0(X17)) & ~(v1_xboole_0(k3_gate_1(X4,X8)) & ~v1_xboole_0(X16)) & ~(v1_xboole_0(k3_gate_1(X3,X7)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(k3_gate_1(X2,X6)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(k3_gate_1(X1,X5)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k11_gate_1(X4,X8,X11))) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k11_gate_1(X3,X7,X10))) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k11_gate_1(X2,X6,X9))) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k11_gate_1(X1,X5,X0)))) [input t40_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : ! [X8] : ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ! [X15] : ! [X16] : ! [X17] : ! [X18] : ~(~(~(v1_xboole_0(X12) & ~v1_xboole_0(X18)) & ~(v1_xboole_0(X18) & ~v1_xboole_0(X12))) & ~(v1_xboole_0(X18) & ~v1_xboole_0(k3_gate_1(X12,X17))) & ~(v1_xboole_0(k3_gate_1(X12,X17)) & ~v1_xboole_0(X18)) & ~(v1_xboole_0(k18_gate_1(X0,X13,X14,X15,X16)) & ~v1_xboole_0(X17)) & ~(v1_xboole_0(k3_gate_1(X4,X8)) & ~v1_xboole_0(X16)) & ~(v1_xboole_0(k3_gate_1(X3,X7)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(k3_gate_1(X2,X6)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(k3_gate_1(X1,X5)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k11_gate_1(X4,X8,X11))) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k11_gate_1(X3,X7,X10))) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k11_gate_1(X2,X6,X9))) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k11_gate_1(X1,X5,X0)))) [negated conjecture 1]
3. ! [X0,X1] : k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [input commutativity_k3_gate_1]
4. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((~(~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0) & ~(k18_gate_1(X0,X1,X2,X3,X4) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input d18_gate_1]
11. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
14. ! [X0,X1,X2] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k11_gate_1(X0,X2,X1))) [input fc34_gate_1]
15. ! [X0,X1,X2] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k11_gate_1(X2,X0,X1))) [input fc35_gate_1]
19. ! [X0,X1] : (~v1_xboole_0(X1) => ~v1_xboole_0(k3_gate_1(X0,X1))) [input fc5_gate_1]
20. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k3_gate_1(X0,X1))) [input fc6_gate_1]
23. ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : ~(~(~(v1_xboole_0(X12) & ~v1_xboole_0(X18)) & ~(v1_xboole_0(X18) & ~v1_xboole_0(X12))) & ~(v1_xboole_0(X18) & ~v1_xboole_0(k3_gate_1(X12,X17))) & ~(v1_xboole_0(k3_gate_1(X12,X17)) & ~v1_xboole_0(X18)) & ~(v1_xboole_0(k18_gate_1(X0,X13,X14,X15,X16)) & ~v1_xboole_0(X17)) & ~(v1_xboole_0(k3_gate_1(X4,X8)) & ~v1_xboole_0(X16)) & ~(v1_xboole_0(k3_gate_1(X3,X7)) & ~v1_xboole_0(X15)) & ~(v1_xboole_0(k3_gate_1(X2,X6)) & ~v1_xboole_0(X14)) & ~(v1_xboole_0(k3_gate_1(X1,X5)) & ~v1_xboole_0(X13)) & ~(v1_xboole_0(X12) & ~v1_xboole_0(k11_gate_1(X4,X8,X11))) & ~(v1_xboole_0(X11) & ~v1_xboole_0(k11_gate_1(X3,X7,X10))) & ~(v1_xboole_0(X10) & ~v1_xboole_0(k11_gate_1(X2,X6,X9))) & ~(v1_xboole_0(X9) & ~v1_xboole_0(k11_gate_1(X1,X5,X0)))) [flattening 2]
24. ! [X0,X1,X2,X3,X4] : ((~(~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0) & ~(k18_gate_1(X0,X1,X2,X3,X4) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 4]
25. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (((v1_xboole_0(X12) & ~v1_xboole_0(X18)) | (v1_xboole_0(X18) & ~v1_xboole_0(X12))) & (~v1_xboole_0(X18) | v1_xboole_0(k3_gate_1(X12,X17))) & (~v1_xboole_0(k3_gate_1(X12,X17)) | v1_xboole_0(X18)) & (~v1_xboole_0(k18_gate_1(X0,X13,X14,X15,X16)) | v1_xboole_0(X17)) & (~v1_xboole_0(k3_gate_1(X4,X8)) | v1_xboole_0(X16)) & (~v1_xboole_0(k3_gate_1(X3,X7)) | v1_xboole_0(X15)) & (~v1_xboole_0(k3_gate_1(X2,X6)) | v1_xboole_0(X14)) & (~v1_xboole_0(k3_gate_1(X1,X5)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k11_gate_1(X4,X8,X11))) & (~v1_xboole_0(X11) | v1_xboole_0(k11_gate_1(X3,X7,X10))) & (~v1_xboole_0(X10) | v1_xboole_0(k11_gate_1(X2,X6,X9))) & (~v1_xboole_0(X9) | v1_xboole_0(k11_gate_1(X1,X5,X0)))) [ennf transformation 23]
28. ! [X0,X1,X2,X3,X4] : ((k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | (~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & (k18_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 24]
29. ! [X0,X1] : (~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1)) [ennf transformation 19]
30. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 20]
31. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 30]
32. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X2,X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 15]
33. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X2,X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 32]
36. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X0,X2,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 14]
37. ! [X0,X1,X2] : (~v1_xboole_0(k11_gate_1(X0,X2,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 36]
44. ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (((v1_xboole_0(X12) & ~v1_xboole_0(X18)) | (v1_xboole_0(X18) & ~v1_xboole_0(X12))) & (~v1_xboole_0(X18) | v1_xboole_0(k3_gate_1(X12,X17))) & (~v1_xboole_0(k3_gate_1(X12,X17)) | v1_xboole_0(X18)) & (~v1_xboole_0(k18_gate_1(X0,X13,X14,X15,X16)) | v1_xboole_0(X17)) & (~v1_xboole_0(k3_gate_1(X4,X8)) | v1_xboole_0(X16)) & (~v1_xboole_0(k3_gate_1(X3,X7)) | v1_xboole_0(X15)) & (~v1_xboole_0(k3_gate_1(X2,X6)) | v1_xboole_0(X14)) & (~v1_xboole_0(k3_gate_1(X1,X5)) | v1_xboole_0(X13)) & (~v1_xboole_0(X12) | v1_xboole_0(k11_gate_1(X4,X8,X11))) & (~v1_xboole_0(X11) | v1_xboole_0(k11_gate_1(X3,X7,X10))) & (~v1_xboole_0(X10) | v1_xboole_0(k11_gate_1(X2,X6,X9))) & (~v1_xboole_0(X9) | v1_xboole_0(k11_gate_1(X1,X5,X0)))) => (((v1_xboole_0(sK12) & ~v1_xboole_0(sK18)) | (v1_xboole_0(sK18) & ~v1_xboole_0(sK12))) & (~v1_xboole_0(sK18) | v1_xboole_0(k3_gate_1(sK12,sK17))) & (~v1_xboole_0(k3_gate_1(sK12,sK17)) | v1_xboole_0(sK18)) & (~v1_xboole_0(k18_gate_1(sK0,sK13,sK14,sK15,sK16)) | v1_xboole_0(sK17)) & (~v1_xboole_0(k3_gate_1(sK4,sK8)) | v1_xboole_0(sK16)) & (~v1_xboole_0(k3_gate_1(sK3,sK7)) | v1_xboole_0(sK15)) & (~v1_xboole_0(k3_gate_1(sK2,sK6)) | v1_xboole_0(sK14)) & (~v1_xboole_0(k3_gate_1(sK1,sK5)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k11_gate_1(sK4,sK8,sK11))) & (~v1_xboole_0(sK11) | v1_xboole_0(k11_gate_1(sK3,sK7,sK10))) & (~v1_xboole_0(sK10) | v1_xboole_0(k11_gate_1(sK2,sK6,sK9))) & (~v1_xboole_0(sK9) | v1_xboole_0(k11_gate_1(sK1,sK5,sK0)))) [choice axiom]
45. ((v1_xboole_0(sK12) & ~v1_xboole_0(sK18)) | (v1_xboole_0(sK18) & ~v1_xboole_0(sK12))) & (~v1_xboole_0(sK18) | v1_xboole_0(k3_gate_1(sK12,sK17))) & (~v1_xboole_0(k3_gate_1(sK12,sK17)) | v1_xboole_0(sK18)) & (~v1_xboole_0(k18_gate_1(sK0,sK13,sK14,sK15,sK16)) | v1_xboole_0(sK17)) & (~v1_xboole_0(k3_gate_1(sK4,sK8)) | v1_xboole_0(sK16)) & (~v1_xboole_0(k3_gate_1(sK3,sK7)) | v1_xboole_0(sK15)) & (~v1_xboole_0(k3_gate_1(sK2,sK6)) | v1_xboole_0(sK14)) & (~v1_xboole_0(k3_gate_1(sK1,sK5)) | v1_xboole_0(sK13)) & (~v1_xboole_0(sK12) | v1_xboole_0(k11_gate_1(sK4,sK8,sK11))) & (~v1_xboole_0(sK11) | v1_xboole_0(k11_gate_1(sK3,sK7,sK10))) & (~v1_xboole_0(sK10) | v1_xboole_0(k11_gate_1(sK2,sK6,sK9))) & (~v1_xboole_0(sK9) | v1_xboole_0(k11_gate_1(sK1,sK5,sK0))) [skolemisation 25,44]
50. ~v1_xboole_0(sK9) | v1_xboole_0(k11_gate_1(sK1,sK5,sK0)) [cnf transformation 45] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. ~v1_xboole_0(sK10) | v1_xboole_0(k11_gate_1(sK2,sK6,sK9)) [cnf transformation 45] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. ~v1_xboole_0(sK11) | v1_xboole_0(k11_gate_1(sK3,sK7,sK10)) [cnf transformation 45] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. ~v1_xboole_0(sK12) | v1_xboole_0(k11_gate_1(sK4,sK8,sK11)) [cnf transformation 45] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. ~v1_xboole_0(k3_gate_1(sK1,sK5)) | v1_xboole_0(sK13) [cnf transformation 45] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. ~v1_xboole_0(k3_gate_1(sK2,sK6)) | v1_xboole_0(sK14) [cnf transformation 45] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
56. ~v1_xboole_0(k3_gate_1(sK3,sK7)) | v1_xboole_0(sK15) [cnf transformation 45] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
57. ~v1_xboole_0(k3_gate_1(sK4,sK8)) | v1_xboole_0(sK16) [cnf transformation 45] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. ~v1_xboole_0(k18_gate_1(sK0,sK13,sK14,sK15,sK16)) | v1_xboole_0(sK17) [cnf transformation 45] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. ~v1_xboole_0(k3_gate_1(sK12,sK17)) | v1_xboole_0(sK18) [cnf transformation 45] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
60. ~v1_xboole_0(sK18) | v1_xboole_0(k3_gate_1(sK12,sK17)) [cnf transformation 45] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
61. ~v1_xboole_0(sK18) | ~v1_xboole_0(sK12) [cnf transformation 45] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
64. v1_xboole_0(sK12) | v1_xboole_0(sK18) [cnf transformation 45] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. v1_xboole_0(k1_xboole_0) [cnf transformation 11] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
69. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X0) [cnf transformation 28] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X1) [cnf transformation 28] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X2) [cnf transformation 28] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X3) [cnf transformation 28] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X4) [cnf transformation 28] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
74. k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
75. ~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1) [cnf transformation 29] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
76. v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 31] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
77. ~v1_xboole_0(k11_gate_1(X2,X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 33] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. ~v1_xboole_0(k11_gate_1(X0,X2,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 37] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
86. 1 <=> v1_xboole_0(sK18) [avatar definition]
90. 2 <=> v1_xboole_0(sK12) [avatar definition]
91. ~v1_xboole_0(sK12) <- (~2) [avatar component clause 90] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
92. v1_xboole_0(sK12) <- (2) [avatar component clause 90] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
93. 1 | 2 [avatar split clause 64,90,86]
94. ~2 | ~1 [avatar split clause 61,86,90]
96. 3 <=> v1_xboole_0(k3_gate_1(sK12,sK17)) [avatar definition]
97. ~v1_xboole_0(k3_gate_1(sK12,sK17)) <- (~3) [avatar component clause 96] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
98. v1_xboole_0(k3_gate_1(sK12,sK17)) <- (3) [avatar component clause 96] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
99. 3 | ~1 [avatar split clause 60,86,96]
100. 1 | ~3 [avatar split clause 59,96,86]
102. 4 <=> v1_xboole_0(sK17) [avatar definition]
106. 5 <=> v1_xboole_0(k18_gate_1(sK0,sK13,sK14,sK15,sK16)) [avatar definition]
108. ~v1_xboole_0(k18_gate_1(sK0,sK13,sK14,sK15,sK16)) <- (~5) [avatar component clause 106] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
109. 4 | ~5 [avatar split clause 58,106,102]
111. 6 <=> v1_xboole_0(sK16) [avatar definition]
113. v1_xboole_0(sK16) <- (6) [avatar component clause 111] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
115. 7 <=> v1_xboole_0(k3_gate_1(sK4,sK8)) [avatar definition]
117. ~v1_xboole_0(k3_gate_1(sK4,sK8)) <- (~7) [avatar component clause 115] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
118. 6 | ~7 [avatar split clause 57,115,111]
120. 8 <=> v1_xboole_0(sK15) [avatar definition]
124. 9 <=> v1_xboole_0(k3_gate_1(sK3,sK7)) [avatar definition]
126. ~v1_xboole_0(k3_gate_1(sK3,sK7)) <- (~9) [avatar component clause 124] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
127. 8 | ~9 [avatar split clause 56,124,120]
129. 10 <=> v1_xboole_0(sK14) [avatar definition]
133. 11 <=> v1_xboole_0(k3_gate_1(sK2,sK6)) [avatar definition]
135. ~v1_xboole_0(k3_gate_1(sK2,sK6)) <- (~11) [avatar component clause 133] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
136. 10 | ~11 [avatar split clause 55,133,129]
138. 12 <=> v1_xboole_0(sK13) [avatar definition]
142. 13 <=> v1_xboole_0(k3_gate_1(sK1,sK5)) [avatar definition]
144. ~v1_xboole_0(k3_gate_1(sK1,sK5)) <- (~13) [avatar component clause 142] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
145. 12 | ~13 [avatar split clause 54,142,138]
147. 14 <=> v1_xboole_0(k11_gate_1(sK4,sK8,sK11)) [avatar definition]
149. v1_xboole_0(k11_gate_1(sK4,sK8,sK11)) <- (14) [avatar component clause 147] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
150. 14 | ~2 [avatar split clause 53,90,147]
152. 15 <=> v1_xboole_0(k11_gate_1(sK3,sK7,sK10)) [avatar definition]
154. v1_xboole_0(k11_gate_1(sK3,sK7,sK10)) <- (15) [avatar component clause 152] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
156. 16 <=> v1_xboole_0(sK11) [avatar definition]
159. 15 | ~16 [avatar split clause 52,156,152]
161. 17 <=> v1_xboole_0(k11_gate_1(sK2,sK6,sK9)) [avatar definition]
163. v1_xboole_0(k11_gate_1(sK2,sK6,sK9)) <- (17) [avatar component clause 161] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
165. 18 <=> v1_xboole_0(sK10) [avatar definition]
168. 17 | ~18 [avatar split clause 51,165,161]
170. 19 <=> v1_xboole_0(k11_gate_1(sK1,sK5,sK0)) [avatar definition]
172. v1_xboole_0(k11_gate_1(sK1,sK5,sK0)) <- (19) [avatar component clause 170] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
174. 20 <=> v1_xboole_0(sK9) [avatar definition]
176. ~v1_xboole_0(sK9) <- (~20) [avatar component clause 174] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
177. 19 | ~20 [avatar split clause 50,174,170]
179. 21 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
181. v1_xboole_0(k1_xboole_0) <- (21) [avatar component clause 179] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
182. 21 [avatar split clause 65,179]
194. ~v1_xboole_0(k3_gate_1(X1,X0)) | v1_xboole_0(X1) [superposition 75,74] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
197. ~v1_xboole_0(sK5) | ~v1_xboole_0(sK1) <- (~13) [resolution 76,144] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
198. ~v1_xboole_0(sK6) | ~v1_xboole_0(sK2) <- (~11) [resolution 76,135] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
199. ~v1_xboole_0(sK7) | ~v1_xboole_0(sK3) <- (~9) [resolution 76,126] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
201. ~v1_xboole_0(sK17) | ~v1_xboole_0(sK12) <- (~3) [resolution 76,97] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
205. 24 <=> v1_xboole_0(sK1) [avatar definition]
209. 25 <=> v1_xboole_0(sK5) [avatar definition]
212. ~24 | ~25 | 13 [avatar split clause 197,142,209,205]
214. 26 <=> v1_xboole_0(sK2) [avatar definition]
216. ~v1_xboole_0(sK2) <- (~26) [avatar component clause 214] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
218. 27 <=> v1_xboole_0(sK6) [avatar definition]
221. ~26 | ~27 | 11 [avatar split clause 198,133,218,214]
223. 28 <=> v1_xboole_0(sK3) [avatar definition]
227. 29 <=> v1_xboole_0(sK7) [avatar definition]
230. ~28 | ~29 | 9 [avatar split clause 199,124,227,223]
232. 30 <=> v1_xboole_0(sK4) [avatar definition]
233. v1_xboole_0(sK4) <- (30) [avatar component clause 232] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
236. 31 <=> v1_xboole_0(sK8) [avatar definition]
240. ~v1_xboole_0(sK17) <- (2, ~3) [subsumption resolution 201,92] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
241. ~4 | ~2 | 3 [avatar split clause 240,96,90,102]
250. v1_xboole_0(sK9) | v1_xboole_0(sK6) <- (17) [resolution 163,77] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
251. v1_xboole_0(sK6) <- (17, ~20) [subsumption resolution 250,176] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
252. 27 | ~17 | 20 [avatar split clause 251,174,161,218]
259. v1_xboole_0(sK11) | v1_xboole_0(sK8) <- (14) [resolution 149,77] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
260. v1_xboole_0(sK10) | v1_xboole_0(sK7) <- (15) [resolution 154,77] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
270. v1_xboole_0(sK11) | v1_xboole_0(sK4) <- (14) [resolution 79,149] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
271. v1_xboole_0(sK10) | v1_xboole_0(sK3) <- (15) [resolution 79,154] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
277. 28 | 18 | ~15 [avatar split clause 271,152,165,223]
278. 29 | 18 | ~15 [avatar split clause 260,152,165,227]
280. 30 | 16 | ~14 [avatar split clause 270,147,156,232]
281. 31 | 16 | ~14 [avatar split clause 259,147,156,236]
297. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK0) <- (~5) [superposition 108,69] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
298. ~v1_xboole_0(sK0) <- (~5, 21) [subsumption resolution 297,181] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
300. 32 <=> v1_xboole_0(sK0) [avatar definition]
302. ~v1_xboole_0(sK0) <- (~32) [avatar component clause 300] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
303. ~32 | 5 | ~21 [avatar split clause 298,179,106,300]
304. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK13) <- (~5) [superposition 108,70] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
305. ~v1_xboole_0(sK13) <- (~5, 21) [subsumption resolution 304,181] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
306. ~12 | 5 | ~21 [avatar split clause 305,179,106,138]
311. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK14) <- (~5) [superposition 108,71] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
312. ~v1_xboole_0(sK14) <- (~5, 21) [subsumption resolution 311,181] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
313. ~10 | 5 | ~21 [avatar split clause 312,179,106,129]
314. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK15) <- (~5) [superposition 108,72] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
315. ~v1_xboole_0(sK15) <- (~5, 21) [subsumption resolution 314,181] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
316. ~8 | 5 | ~21 [avatar split clause 315,179,106,120]
317. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK16) <- (~5) [superposition 108,73] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
318. ~v1_xboole_0(sK16) <- (~5, 21) [subsumption resolution 317,181] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
319. $false <- (~5, 6, 21) [subsumption resolution 318,113] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
320. 5 | ~6 | ~21 [avatar contradiction clause 319]
322. v1_xboole_0(sK12) <- (3) [resolution 98,194] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
325. $false <- (~2, 3) [subsumption resolution 322,91] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
326. 2 | ~3 [avatar contradiction clause 325]
328. ~v1_xboole_0(sK8) | ~v1_xboole_0(sK4) <- (~7) [resolution 117,76] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
329. ~v1_xboole_0(sK8) <- (~7, 30) [subsumption resolution 328,233] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
330. ~31 | 7 | ~30 [avatar split clause 329,232,115,236]
334. v1_xboole_0(sK9) | v1_xboole_0(sK2) <- (17) [resolution 163,79] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
337. v1_xboole_0(sK9) <- (17, ~26) [subsumption resolution 334,216] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
338. 20 | ~17 | 26 [avatar split clause 337,214,161,174]
340. v1_xboole_0(sK0) | v1_xboole_0(sK1) <- (19) [resolution 172,79] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
342. v1_xboole_0(sK0) | v1_xboole_0(sK5) <- (19) [resolution 172,77] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
344. v1_xboole_0(sK1) <- (19, ~32) [subsumption resolution 340,302] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
348. v1_xboole_0(sK5) <- (19, ~32) [subsumption resolution 342,302] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
350. 24 | ~19 | 32 [avatar split clause 344,300,170,205]
351. 25 | ~19 | 32 [avatar split clause 348,300,170,209]
352. $false [avatar sat refutation 93,94,99,100,109,118,127,136,145,150,159,168,177,182,212,221,230,241,252,277,278,280,281,303,306,313,316,320,326,330,338,350,351] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:271,thDist:-271}
% SZS output end Proof for gate_1__t40_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 21
% 
% >>> Saturation
% Initial clauses: 35
% Generated clauses: 208
% Active clauses: 85
% Passive clauses: 95
% Final active clauses: 46
% Final passive clauses: 3
% Main loop iterations started: 86
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 22
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 49
% Unit resulting resolution: 4
% Forward superposition: 4
% Backward superposition: 7
% 
% >>> AVATAR
% Split clauses: 26
% Split components: 52
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 55
% SAT solver unit clauses: 3
% SAT solver binary clauses: 18
% 
% Memory used [KB]: 5245
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

