i: [1,46,0,3,1,1,0,0] t47_zfmisc_1
a: [3,75,0,3,1,46]
s: 75
i: [1,45,0,3,1,1,0,0] t47_zfmisc_1
a: [3,80,0,3,1,45]
s: 80
i: [1,48,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,85,0,2,1,48]
s: 85
i: [1,49,0,2,1,0,0,255] fc1_xboole_0
a: [3,90,0,2,1,49]
s: 90
i: [1,65,0,2,1,0,0,2] rc2_xboole_0
a: [3,100,0,2,1,65]
s: 100
i: [1,66,0,2,1,0,0,2] rc1_xboole_0
a: [3,105,0,2,1,66]
s: 105
i: [1,47,0,9,1,1,0,0] t47_zfmisc_1
a: [3,70,0,9,1,47]
s: 70
i: [1,59,0,3,1,0,0,255] reflexivity_r1_tarski
s: 59
i: [1,50,0,3,1,0,0,255] d2_xboole_0
a: [3,95,0,3,1,50]
s: 95
r: 90
i: [1,60,0,4,1,0,0,1] fc3_xboole_0
s: 60
i: [1,51,0,5,1,0,0,255] t2_boole
s: 51
i: [1,52,0,5,2,0,0,255] t6_boole
s: 52
i: [1,55,0,5,2,0,0,1] t7_boole
s: 55
d: [2,110,1,2,1,1,40,55,75]
a: [3,115,1,2,1,110]
s: 115
i: [1,61,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 61
d: [2,108,0,5,1,1,22,51,95]
s: 108
d: [2,109,0,5,2,1,22,52,95]
s: 109
i: [1,64,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 64
d: [2,121,1,3,1,1,40,64,75]
a: [3,126,1,3,1,121]
s: 126
i: [1,54,0,7,3,0,0,2] t8_boole
s: 54
i: [1,62,0,7,1,0,0,1] commutativity_k2_tarski
s: 62
d: [2,122,1,3,1,1,40,64,80]
a: [3,131,1,3,1,122]
s: 131
i: [1,63,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 63
d: [2,137,1,9,1,1,44,70,63]
a: [3,144,1,9,1,137]
s: 144
i: [1,53,0,8,2,0,0,255] t28_xboole_1
s: 53
i: [1,56,0,8,2,0,0,255] t32_zfmisc_1
s: 56
d: [2,150,1,12,2,1,44,70,53]
d: [2,155,1,5,1,1,19,150]
a: [3,159,1,5,1,155]
s: 159
i: [1,57,0,8,2,0,0,255] t32_zfmisc_1
s: 57
d: [2,154,2,10,2,1,44,144,53]
a: [3,164,2,5,1,154]
s: 164
i: [1,58,0,11,3,0,0,255] t32_zfmisc_1
s: 58
d: [2,176,2,0,0,3,52,80,75,159,58]
e: 176
d: [2,179,2,6,2,1,40,58,159]
d: [2,183,2,3,1,2,21,179,75]
d: [2,184,2,0,0,3,21,183,80]
e: 184
f: 184,155,45,46,47
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t47_zfmisc_1
% SZS output start Proof for zfmisc_1__t47_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ((r2_hidden(X2,X1) & r2_hidden(X0,X1)) => k2_tarski(X0,X2) = k3_xboole_0(k2_tarski(X0,X2),X1)) [input t47_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ((r2_hidden(X2,X1) & r2_hidden(X0,X1)) => k2_tarski(X0,X2) = k3_xboole_0(k2_tarski(X0,X2),X1)) [negated conjecture 1]
21. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
22. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),X2) <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [input t32_zfmisc_1]
23. ~! [X0,X1,X2] : ((r2_hidden(X2,X1) & r2_hidden(X0,X1)) => k2_tarski(X0,X2) = k3_xboole_0(k2_tarski(X0,X2),X1)) [flattening 2]
24. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 21]
27. ! [X0,X1,X2] : (r1_tarski(k2_tarski(X0,X1),X2) <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [flattening 22]
30. ? [X0,X1,X2] : (k2_tarski(X0,X2) != k3_xboole_0(k2_tarski(X0,X2),X1) & (r2_hidden(X2,X1) & r2_hidden(X0,X1))) [ennf transformation 23]
31. ? [X0,X1,X2] : (k2_tarski(X0,X2) != k3_xboole_0(k2_tarski(X0,X2),X1) & r2_hidden(X2,X1) & r2_hidden(X0,X1)) [flattening 30]
33. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 24]
37. ? [X0,X1,X2] : (k2_tarski(X0,X2) != k3_xboole_0(k2_tarski(X0,X2),X1) & r2_hidden(X2,X1) & r2_hidden(X0,X1)) => (k2_tarski(sK0,sK2) != k3_xboole_0(k2_tarski(sK0,sK2),sK1) & r2_hidden(sK2,sK1) & r2_hidden(sK0,sK1)) [choice axiom]
38. k2_tarski(sK0,sK2) != k3_xboole_0(k2_tarski(sK0,sK2),sK1) & r2_hidden(sK2,sK1) & r2_hidden(sK0,sK1) [skolemisation 31,37]
39. ! [X0,X1,X2] : ((r1_tarski(k2_tarski(X0,X1),X2) | (~r2_hidden(X1,X2) | ~r2_hidden(X0,X2))) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | ~r1_tarski(k2_tarski(X0,X1),X2))) [nnf transformation 27]
40. ! [X0,X1,X2] : ((r1_tarski(k2_tarski(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2)) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | ~r1_tarski(k2_tarski(X0,X1),X2))) [flattening 39]
45. r2_hidden(sK0,sK1) [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. r2_hidden(sK2,sK1) [cnf transformation 38] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
47. k2_tarski(sK0,sK2) != k3_xboole_0(k2_tarski(sK0,sK2),sK1) [cnf transformation 38] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
53. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 33] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
58. r1_tarski(k2_tarski(X0,X1),X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) [cnf transformation 40] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
68. 1 <=> k2_tarski(sK0,sK2) = k3_xboole_0(k2_tarski(sK0,sK2),sK1) [avatar definition]
70. k2_tarski(sK0,sK2) != k3_xboole_0(k2_tarski(sK0,sK2),sK1) <- (~1) [avatar component clause 68] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. ~1 [avatar split clause 47,68]
73. 2 <=> r2_hidden(sK2,sK1) [avatar definition]
75. r2_hidden(sK2,sK1) <- (2) [avatar component clause 73] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. 2 [avatar split clause 46,73]
78. 3 <=> r2_hidden(sK0,sK1) [avatar definition]
80. r2_hidden(sK0,sK1) <- (3) [avatar component clause 78] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
81. 3 [avatar split clause 45,78]
150. k2_tarski(sK0,sK2) != k2_tarski(sK0,sK2) | ~r1_tarski(k2_tarski(sK0,sK2),sK1) <- (~1) [superposition 70,53] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
155. ~r1_tarski(k2_tarski(sK0,sK2),sK1) <- (~1) [trivial inequality removal 150] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
157. 13 <=> r1_tarski(k2_tarski(sK0,sK2),sK1) [avatar definition]
159. ~r1_tarski(k2_tarski(sK0,sK2),sK1) <- (~13) [avatar component clause 157] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
160. ~13 | 1 [avatar split clause 155,68,157]
179. ~r2_hidden(sK2,sK1) | ~r2_hidden(sK0,sK1) <- (~13) [resolution 58,159] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
183. ~r2_hidden(sK0,sK1) <- (2, ~13) [subsumption resolution 179,75] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
184. $false <- (2, 3, ~13) [subsumption resolution 183,80] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
185. ~2 | ~3 | 13 [avatar contradiction clause 184]
186. $false [avatar sat refutation 71,76,81,160,185] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:75,thDist:-75}
% SZS output end Proof for zfmisc_1__t47_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 22
% Generated clauses: 77
% Active clauses: 28
% Passive clauses: 46
% Final active clauses: 27
% Final passive clauses: 16
% Main loop iterations started: 30
% 
% >>> Simplifying Inferences
% Trivial inequalities: 1
% Fw subsumption resolutions: 2
% Fw demodulations: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 1
% Forward superposition: 11
% Backward superposition: 13
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 8
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

