i: [1,46,0,7,2,1,0,0] l1_zfmisc_1
a: [3,74,0,3,1,46]
s: 74
a: [3,71,0,4,1,46]
s: 71
i: [1,47,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,81,0,2,1,47]
s: 81
i: [1,48,0,2,1,0,0,255] fc1_xboole_0
a: [3,86,0,2,1,48]
s: 86
i: [1,63,0,2,1,0,0,2] rc2_xboole_0
a: [3,96,0,2,1,63]
s: 96
i: [1,64,0,2,1,0,0,2] rc1_xboole_0
a: [3,101,0,2,1,64]
s: 101
i: [1,50,0,3,1,0,0,1] fc2_xboole_0
s: 50
i: [1,61,0,3,1,0,0,1] reflexivity_r1_tarski
s: 61
i: [1,49,0,3,1,0,0,255] d2_xboole_0
a: [3,91,0,3,1,49]
s: 91
r: 81
i: [1,66,0,4,1,0,0,1] d1_tarski
s: 66
i: [1,51,0,5,2,0,0,255] t6_boole
s: 51
i: [1,60,0,5,2,0,0,1] t7_boole
s: 60
i: [1,62,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 62
d: [2,106,1,2,1,1,40,60,74]
a: [3,111,1,2,1,106]
s: 111
i: [1,59,0,7,3,0,0,2] t8_boole
s: 59
i: [1,67,0,7,2,0,0,1] d1_tarski
s: 67
d: [2,113,1,3,1,1,40,62,74]
a: [3,118,1,3,1,113]
s: 118
i: [1,53,0,8,2,0,0,1] d3_tarski
s: 53
d: [2,114,1,4,1,0,40,62,66]
s: 114
i: [1,54,0,8,2,0,0,1] d3_tarski
s: 54
i: [1,52,0,9,3,0,0,1] d3_tarski
s: 52
d: [2,126,1,6,2,1,40,52,74]
s: 126
i: [1,57,0,14,3,0,0,1] d1_tarski
s: 57
d: [2,122,1,5,2,0,40,53,60]
s: 122
i: [1,58,0,14,3,0,0,1] d1_tarski
s: 58
d: [2,104,1,7,3,0,44,51,51]
s: 104
d: [2,127,1,7,2,0,40,52,66]
s: 127
d: [2,121,1,8,2,0,40,53,62]
s: 121
d: [2,123,1,10,2,0,40,53,67]
s: 123
d: [2,128,1,11,3,0,40,52,53]
s: 128
d: [2,132,1,11,3,0,40,57,60]
s: 132
d: [2,139,2,11,3,0,44,54,123]
d: [2,142,2,7,2,0,18,139]
s: 142
d: [2,157,3,0,0,2,52,74,71,142]
e: 157
d: [2,158,3,3,1,1,40,142,71]
d: [2,161,3,0,0,2,21,158,74]
e: 161
r: 71
r: 126
r: 74
a: [3,75,0,3,1,46]
s: 75
a: [3,70,0,4,1,46]
s: 70
d: [2,163,2,0,0,2,52,75,70,127]
e: 163
d: [2,164,1,0,0,2,52,75,66,70,52]
e: 164
d: [2,165,2,3,1,1,40,70,127]
d: [2,168,2,0,0,2,21,165,75]
e: 168
i: [1,45,0,7,2,1,0,0] l1_zfmisc_1
f: 168,161,45,46
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__l1_zfmisc_1
% SZS output start Proof for zfmisc_1__l1_zfmisc_1
1. ! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [input l1_zfmisc_1]
2. ~! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [negated conjecture 1]
4. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
6. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
18. ~! [X0,X1] : (r1_tarski(k1_tarski(X0),X1) <=> r2_hidden(X0,X1)) [flattening 2]
19. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 6]
20. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 4]
24. ? [X0,X1] : (r1_tarski(k1_tarski(X0),X1) <~> r2_hidden(X0,X1)) [ennf transformation 18]
26. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 19]
30. ? [X0,X1] : ((~r2_hidden(X0,X1) | ~r1_tarski(k1_tarski(X0),X1)) & (r2_hidden(X0,X1) | r1_tarski(k1_tarski(X0),X1))) [nnf transformation 24]
31. ? [X0,X1] : ((~r2_hidden(X0,X1) | ~r1_tarski(k1_tarski(X0),X1)) & (r2_hidden(X0,X1) | r1_tarski(k1_tarski(X0),X1))) => ((~r2_hidden(sK0,sK1) | ~r1_tarski(k1_tarski(sK0),sK1)) & (r2_hidden(sK0,sK1) | r1_tarski(k1_tarski(sK0),sK1))) [choice axiom]
32. (~r2_hidden(sK0,sK1) | ~r1_tarski(k1_tarski(sK0),sK1)) & (r2_hidden(sK0,sK1) | r1_tarski(k1_tarski(sK0),sK1)) [skolemisation 30,31]
33. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 26]
34. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 33]
35. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) [choice axiom]
36. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK2(X0,X1),X1) & r2_hidden(sK2(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 34,35]
37. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 20]
38. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 37]
39. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
40. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 38,39]
45. r2_hidden(sK0,sK1) | r1_tarski(k1_tarski(sK0),sK1) [cnf transformation 32] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
46. ~r2_hidden(sK0,sK1) | ~r1_tarski(k1_tarski(sK0),sK1) [cnf transformation 32] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
52. ~r2_hidden(X3,X0) | r2_hidden(X3,X1) | ~r1_tarski(X0,X1) [cnf transformation 36] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
53. r2_hidden(sK2(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
54. ~r2_hidden(sK2(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
55. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 40] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
56. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 40] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
65. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 56] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:2,thDist:-2}
66. r2_hidden(X3,k1_tarski(X3)) [equality resolution 65] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
67. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 55] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
69. 1 <=> r1_tarski(k1_tarski(sK0),sK1) [avatar definition]
70. r1_tarski(k1_tarski(sK0),sK1) <- (1) [avatar component clause 69] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. ~r1_tarski(k1_tarski(sK0),sK1) <- (~1) [avatar component clause 69] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
73. 2 <=> r2_hidden(sK0,sK1) [avatar definition]
74. r2_hidden(sK0,sK1) <- (2) [avatar component clause 73] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. ~r2_hidden(sK0,sK1) <- (~2) [avatar component clause 73] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. ~1 | ~2 [avatar split clause 46,73,69]
77. 1 | 2 [avatar split clause 45,73,69]
123. sK2(k1_tarski(X4),X5) = X4 | r1_tarski(k1_tarski(X4),X5) [resolution 53,67] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
127. ~r1_tarski(k1_tarski(X1),X2) | r2_hidden(X1,X2) [resolution 52,66] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
139. ~r2_hidden(X2,X3) | r1_tarski(k1_tarski(X2),X3) | r1_tarski(k1_tarski(X2),X3) [superposition 54,123] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:6,thDist:-6}
142. r1_tarski(k1_tarski(X2),X3) | ~r2_hidden(X2,X3) [duplicate literal removal 139] {a:2,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
158. ~r2_hidden(sK0,sK1) <- (~1) [resolution 142,71] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
161. $false <- (~1, 2) [subsumption resolution 158,74] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
162. 1 | ~2 [avatar contradiction clause 161]
165. r2_hidden(sK0,sK1) <- (1) [resolution 70,127] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
168. $false <- (1, ~2) [subsumption resolution 165,75] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
169. ~1 | 2 [avatar contradiction clause 168]
170. $false [avatar sat refutation 76,77,162,169] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:70,thDist:-70}
% SZS output end Proof for zfmisc_1__l1_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 20
% Generated clauses: 82
% Active clauses: 33
% Passive clauses: 42
% Final active clauses: 29
% Final passive clauses: 8
% Main loop iterations started: 34
% 
% >>> Simplifying Inferences
% Duplicate literals: 7
% Trivial inequalities: 1
% Fw subsumption resolutions: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 30
% Unit resulting resolution: 3
% Forward superposition: 1
% Backward superposition: 5
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 5
% SAT solver binary clauses: 9
% 
% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

