i: [1,75,0,4,2,1,0,0] t11_conmetr1
a: [3,106,0,2,1,75]
s: 106
a: [3,109,0,2,1,75]
s: 109
i: [1,73,0,2,1,1,0,0] t11_conmetr1
a: [3,116,0,2,1,73]
s: 116
i: [1,72,0,2,1,1,0,0] t11_conmetr1
a: [3,121,0,2,1,72]
s: 121
i: [1,71,0,2,1,1,0,0] t11_conmetr1
a: [3,126,0,2,1,71]
s: 126
i: [1,70,0,2,1,1,0,0] t11_conmetr1
a: [3,131,0,2,1,70]
s: 131
i: [1,101,0,2,1,0,0,1] existence_l1_analoaf
a: [3,136,0,2,1,101]
s: 136
i: [1,102,0,2,1,0,0,2] existence_l1_struct_0
a: [3,141,0,2,1,102]
s: 141
i: [1,76,0,4,2,0,0,1] dt_l1_analoaf
s: 76
i: [1,77,0,6,3,0,0,255] cc7_struct_0
s: 77
i: [1,78,0,6,3,0,0,255] cc5_struct_0
s: 78
d: [2,143,1,2,1,1,40,76,116]
a: [3,148,1,2,1,143]
s: 148
i: [1,80,0,6,3,0,0,255] cc17_struct_0
s: 80
d: [2,144,1,2,1,1,40,76,136]
a: [3,153,1,2,1,144]
s: 153
i: [1,81,0,6,3,0,0,255] cc2_struct_0
s: 81
i: [1,82,0,6,3,0,0,255] cc1_struct_0
s: 82
d: [2,155,1,4,2,1,40,77,141]
a: [3,159,1,2,1,155]
s: 159
i: [1,85,0,6,3,0,0,255] cc4_struct_0
s: 85
d: [2,166,2,4,2,1,40,153,77]
a: [3,170,2,2,1,166]
s: 170
i: [1,86,0,6,3,0,0,255] cc6_struct_0
s: 86
i: [1,83,0,7,3,0,0,255] cc8_struct_0
s: 83
i: [1,87,0,7,3,0,0,255] cc9_struct_0
s: 87
i: [1,88,0,7,3,0,0,255] cc11_struct_0
s: 88
i: [1,89,0,7,3,0,0,255] cc11_struct_0
s: 89
i: [1,90,0,9,4,0,0,255] cc10_struct_0
s: 90
i: [1,91,0,12,6,0,0,2] t9_conmetr1
s: 91
i: [1,92,0,12,6,0,0,2] t9_conmetr1
s: 92
i: [1,93,0,12,6,0,0,2] t10_conmetr1
s: 93
i: [1,94,0,12,6,0,0,2] t10_conmetr1
s: 94
i: [1,95,0,12,6,0,0,1] t10_aff_2
s: 95
i: [1,96,0,12,6,0,0,1] t10_aff_2
s: 96
i: [1,98,0,12,6,0,0,1] t1_conmetr1
s: 98
r: 106
r: 109
a: [3,105,0,2,1,75]
s: 105
r: 105
s: 106
s: 109
i: [1,99,0,12,6,0,0,1] t1_conmetr1
s: 99
r: 106
r: 109
s: 105
a: [3,110,0,2,1,75]
s: 110
d: [2,182,1,10,5,1,40,98,109]
d: [2,183,1,8,4,2,21,182,116]
d: [2,184,1,6,3,3,21,183,121]
d: [2,185,1,4,2,4,21,184,126]
d: [2,186,1,2,1,5,21,185,131]
a: [3,190,1,2,1,186]
s: 190
i: [1,97,0,14,7,0,0,1] t10_aff_2
s: 97
i: [1,100,0,14,7,0,0,1] t1_conmetr1
s: 100
d: [2,192,1,10,5,1,40,105,96]
d: [2,194,1,8,4,2,21,192,116]
d: [2,195,1,6,3,3,21,194,121]
d: [2,196,1,4,2,4,21,195,126]
d: [2,197,1,2,1,5,21,196,131]
a: [3,201,1,2,1,197]
s: 201
d: [2,193,1,10,5,1,40,105,95]
d: [2,203,1,8,4,2,21,193,116]
d: [2,204,1,6,3,3,21,203,121]
d: [2,205,1,4,2,4,21,204,126]
d: [2,206,1,2,1,5,21,205,131]
a: [3,210,1,2,1,206]
s: 210
r: 190
r: 105
r: 110
r: 210
s: 106
s: 109
s: 190
a: [3,209,1,2,1,206]
s: 209
d: [2,218,1,10,5,1,40,99,109]
d: [2,219,1,8,4,2,21,218,116]
d: [2,220,1,6,3,3,21,219,121]
d: [2,221,1,4,2,4,21,220,126]
d: [2,222,1,2,1,5,21,221,131]
a: [3,226,1,2,1,222]
s: 226
d: [2,291,2,0,0,6,52,131,126,121,116,209,226,94]
e: 291
d: [2,292,2,10,5,1,40,226,94]
d: [2,294,2,8,4,2,21,292,209]
d: [2,295,2,6,3,3,21,294,116]
d: [2,296,2,4,2,4,21,295,121]
d: [2,297,2,2,1,5,21,296,126]
d: [2,298,2,0,0,6,21,297,131]
e: 298
i: [1,74,0,4,2,1,0,0] t11_conmetr1
d: [2,228,1,10,5,1,40,105,96]
d: [2,230,1,8,4,2,21,228,116]
d: [2,231,1,6,3,3,21,230,121]
d: [2,232,1,4,2,4,21,231,126]
d: [2,233,1,2,1,5,21,232,131]
d: [2,229,1,10,5,1,40,105,95]
d: [2,235,1,8,4,2,21,229,116]
d: [2,236,1,6,3,3,21,235,121]
d: [2,237,1,4,2,4,21,236,126]
d: [2,238,1,2,1,5,21,237,131]
d: [2,240,2,10,5,1,40,190,92]
d: [2,241,2,8,4,2,21,240,116]
d: [2,242,2,6,3,3,21,241,121]
d: [2,243,2,4,2,4,21,242,126]
d: [2,244,2,2,1,5,21,243,131]
d: [2,246,2,12,6,1,40,100,190]
d: [2,247,2,10,5,2,21,246,110]
d: [2,248,2,8,4,3,21,247,116]
d: [2,249,2,6,3,4,21,248,121]
d: [2,250,2,4,2,5,21,249,126]
d: [2,251,2,2,1,6,21,250,131]
d: [2,255,2,10,5,1,40,210,93]
d: [2,256,2,8,4,2,21,255,116]
d: [2,257,2,6,3,3,21,256,121]
d: [2,258,2,4,2,4,21,257,126]
d: [2,259,2,2,1,5,21,258,131]
d: [2,254,2,10,5,1,40,201,91]
d: [2,261,2,8,4,2,21,254,116]
d: [2,262,2,6,3,3,21,261,121]
d: [2,263,2,4,2,4,21,262,126]
d: [2,264,2,2,1,5,21,263,131]
d: [2,253,2,12,6,1,40,201,97]
d: [2,272,2,10,5,2,21,253,116]
d: [2,273,2,8,4,3,21,272,121]
d: [2,274,2,6,3,4,21,273,126]
d: [2,275,2,4,2,5,21,274,131]
d: [2,277,1,10,5,1,40,109,99]
d: [2,279,1,8,4,2,21,277,116]
d: [2,280,1,6,3,3,21,279,121]
d: [2,281,1,4,2,4,21,280,126]
d: [2,282,1,2,1,5,21,281,131]
d: [2,278,1,10,5,1,40,109,98]
d: [2,284,1,8,4,2,21,278,116]
d: [2,285,1,6,3,3,21,284,121]
d: [2,286,1,4,2,4,21,285,126]
d: [2,287,1,2,1,5,21,286,131]
f: 298,287,282,275,264,259,251,244,238,233,70,71,72,73,74,75
% Refutation found. Thanks to Tanya!
% SZS status Theorem for conmetr1__t11_conmetr1
% SZS output start Proof for conmetr1__t11_conmetr1
1. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v1_aff_2(X0) <=> v4_conmetr1(X0))) [input t11_conmetr1]
2. ~! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v1_aff_2(X0) <=> v4_conmetr1(X0))) [negated conjecture 1]
19. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v1_aff_2(X0) <=> (v13_aff_2(X0) & v2_aff_2(X0)))) [input t10_aff_2]
20. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v2_aff_2(X0) <=> v6_conmetr1(X0))) [input t10_conmetr1]
21. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v4_conmetr1(X0) <=> (v6_conmetr1(X0) & v5_conmetr1(X0)))) [input t1_conmetr1]
22. ! [X0] : ((l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (v13_aff_2(X0) <=> v5_conmetr1(X0))) [input t9_conmetr1]
23. ? [X0] : ((v1_aff_2(X0) <~> v4_conmetr1(X0)) & (l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
24. ? [X0] : ((v1_aff_2(X0) <~> v4_conmetr1(X0)) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) [flattening 23]
48. ! [X0] : ((v13_aff_2(X0) <=> v5_conmetr1(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 22]
49. ! [X0] : ((v13_aff_2(X0) <=> v5_conmetr1(X0)) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 48]
50. ! [X0] : ((v2_aff_2(X0) <=> v6_conmetr1(X0)) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 20]
51. ! [X0] : ((v2_aff_2(X0) <=> v6_conmetr1(X0)) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 50]
52. ! [X0] : ((v1_aff_2(X0) <=> (v13_aff_2(X0) & v2_aff_2(X0))) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 19]
53. ! [X0] : ((v1_aff_2(X0) <=> (v13_aff_2(X0) & v2_aff_2(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 52]
54. ! [X0] : ((v4_conmetr1(X0) <=> (v6_conmetr1(X0) & v5_conmetr1(X0))) | (~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 21]
55. ! [X0] : ((v4_conmetr1(X0) <=> (v6_conmetr1(X0) & v5_conmetr1(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 54]
56. ? [X0] : (((~v4_conmetr1(X0) | ~v1_aff_2(X0)) & (v4_conmetr1(X0) | v1_aff_2(X0))) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) [nnf transformation 24]
57. ? [X0] : ((~v4_conmetr1(X0) | ~v1_aff_2(X0)) & (v4_conmetr1(X0) | v1_aff_2(X0)) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) [flattening 56]
58. ? [X0] : ((~v4_conmetr1(X0) | ~v1_aff_2(X0)) & (v4_conmetr1(X0) | v1_aff_2(X0)) & l1_analoaf(X0) & v2_diraf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => ((~v4_conmetr1(sK0) | ~v1_aff_2(sK0)) & (v4_conmetr1(sK0) | v1_aff_2(sK0)) & l1_analoaf(sK0) & v2_diraf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
59. (~v4_conmetr1(sK0) | ~v1_aff_2(sK0)) & (v4_conmetr1(sK0) | v1_aff_2(sK0)) & l1_analoaf(sK0) & v2_diraf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0) [skolemisation 57,58]
60. ! [X0] : (((v13_aff_2(X0) | ~v5_conmetr1(X0)) & (v5_conmetr1(X0) | ~v13_aff_2(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [nnf transformation 49]
61. ! [X0] : (((v2_aff_2(X0) | ~v6_conmetr1(X0)) & (v6_conmetr1(X0) | ~v2_aff_2(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [nnf transformation 51]
62. ! [X0] : (((v1_aff_2(X0) | (~v13_aff_2(X0) | ~v2_aff_2(X0))) & ((v13_aff_2(X0) & v2_aff_2(X0)) | ~v1_aff_2(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [nnf transformation 53]
63. ! [X0] : (((v1_aff_2(X0) | ~v13_aff_2(X0) | ~v2_aff_2(X0)) & ((v13_aff_2(X0) & v2_aff_2(X0)) | ~v1_aff_2(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 62]
64. ! [X0] : (((v4_conmetr1(X0) | (~v6_conmetr1(X0) | ~v5_conmetr1(X0))) & ((v6_conmetr1(X0) & v5_conmetr1(X0)) | ~v4_conmetr1(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [nnf transformation 55]
65. ! [X0] : (((v4_conmetr1(X0) | ~v6_conmetr1(X0) | ~v5_conmetr1(X0)) & ((v6_conmetr1(X0) & v5_conmetr1(X0)) | ~v4_conmetr1(X0))) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 64]
70. ~v7_struct_0(sK0) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
71. v1_diraf(sK0) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
72. v2_diraf(sK0) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
73. l1_analoaf(sK0) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
74. v4_conmetr1(sK0) | v1_aff_2(sK0) [cnf transformation 59] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
75. ~v4_conmetr1(sK0) | ~v1_aff_2(sK0) [cnf transformation 59] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
91. ~v13_aff_2(X0) | v5_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 60] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
92. ~v5_conmetr1(X0) | v13_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 60] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
93. ~v2_aff_2(X0) | v6_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 61] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
94. ~v6_conmetr1(X0) | v2_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 61] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
95. ~v1_aff_2(X0) | v2_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 63] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
96. ~v1_aff_2(X0) | v13_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 63] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
97. ~v13_aff_2(X0) | v1_aff_2(X0) | ~v2_aff_2(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 63] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
98. ~v4_conmetr1(X0) | v5_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 65] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
99. ~v4_conmetr1(X0) | v6_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 65] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
100. ~v5_conmetr1(X0) | ~v6_conmetr1(X0) | v4_conmetr1(X0) | ~l1_analoaf(X0) | ~v2_diraf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 65] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. 1 <=> v1_aff_2(sK0) [avatar definition]
105. v1_aff_2(sK0) <- (1) [avatar component clause 104] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
108. 2 <=> v4_conmetr1(sK0) [avatar definition]
109. v4_conmetr1(sK0) <- (2) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
110. ~v4_conmetr1(sK0) <- (~2) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
111. ~1 | ~2 [avatar split clause 75,108,104]
112. 1 | 2 [avatar split clause 74,108,104]
114. 3 <=> l1_analoaf(sK0) [avatar definition]
116. l1_analoaf(sK0) <- (3) [avatar component clause 114] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
117. 3 [avatar split clause 73,114]
119. 4 <=> v2_diraf(sK0) [avatar definition]
121. v2_diraf(sK0) <- (4) [avatar component clause 119] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
122. 4 [avatar split clause 72,119]
124. 5 <=> v1_diraf(sK0) [avatar definition]
126. v1_diraf(sK0) <- (5) [avatar component clause 124] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
127. 5 [avatar split clause 71,124]
129. 6 <=> v7_struct_0(sK0) [avatar definition]
131. ~v7_struct_0(sK0) <- (~6) [avatar component clause 129] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
132. ~6 [avatar split clause 70,129]
188. 15 <=> v5_conmetr1(sK0) [avatar definition]
190. v5_conmetr1(sK0) <- (15) [avatar component clause 188] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
199. 16 <=> v13_aff_2(sK0) [avatar definition]
201. v13_aff_2(sK0) <- (16) [avatar component clause 199] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
208. 17 <=> v2_aff_2(sK0) [avatar definition]
209. ~v2_aff_2(sK0) <- (~17) [avatar component clause 208] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
210. v2_aff_2(sK0) <- (17) [avatar component clause 208] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
224. 18 <=> v6_conmetr1(sK0) [avatar definition]
226. v6_conmetr1(sK0) <- (18) [avatar component clause 224] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
228. v13_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1) [resolution 105,96] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
229. v2_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1) [resolution 105,95] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
230. v13_aff_2(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1, 3) [subsumption resolution 228,116] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
231. v13_aff_2(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1, 3, 4) [subsumption resolution 230,121] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
232. v13_aff_2(sK0) | v7_struct_0(sK0) <- (1, 3, 4, 5) [subsumption resolution 231,126] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
233. v13_aff_2(sK0) <- (1, 3, 4, 5, ~6) [subsumption resolution 232,131] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
234. 16 | ~1 | ~3 | ~4 | ~5 | 6 [avatar split clause 233,129,124,119,114,104,199]
235. v2_aff_2(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1, 3) [subsumption resolution 229,116] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
236. v2_aff_2(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1, 3, 4) [subsumption resolution 235,121] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
237. v2_aff_2(sK0) | v7_struct_0(sK0) <- (1, 3, 4, 5) [subsumption resolution 236,126] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
238. v2_aff_2(sK0) <- (1, 3, 4, 5, ~6) [subsumption resolution 237,131] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
239. 17 | ~1 | ~3 | ~4 | ~5 | 6 [avatar split clause 238,129,124,119,114,104,208]
240. v13_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (15) [resolution 190,92] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
241. v13_aff_2(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 15) [subsumption resolution 240,116] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
242. v13_aff_2(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 4, 15) [subsumption resolution 241,121] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
243. v13_aff_2(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 15) [subsumption resolution 242,126] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
244. v13_aff_2(sK0) <- (3, 4, 5, ~6, 15) [subsumption resolution 243,131] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
245. 16 | ~3 | ~4 | ~5 | 6 | ~15 [avatar split clause 244,188,129,124,119,114,199]
246. ~v6_conmetr1(sK0) | v4_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (15) [resolution 100,190] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
247. ~v6_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~2, 15) [subsumption resolution 246,110] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
248. ~v6_conmetr1(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~2, 3, 15) [subsumption resolution 247,116] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
249. ~v6_conmetr1(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~2, 3, 4, 15) [subsumption resolution 248,121] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
250. ~v6_conmetr1(sK0) | v7_struct_0(sK0) <- (~2, 3, 4, 5, 15) [subsumption resolution 249,126] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
251. ~v6_conmetr1(sK0) <- (~2, 3, 4, 5, ~6, 15) [subsumption resolution 250,131] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
252. ~18 | 2 | ~3 | ~4 | ~5 | 6 | ~15 [avatar split clause 251,188,129,124,119,114,108,224]
253. v1_aff_2(sK0) | ~v2_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (16) [resolution 201,97] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
254. v5_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (16) [resolution 201,91] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
255. v6_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (17) [resolution 210,93] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
256. v6_conmetr1(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 17) [subsumption resolution 255,116] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
257. v6_conmetr1(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 4, 17) [subsumption resolution 256,121] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
258. v6_conmetr1(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 17) [subsumption resolution 257,126] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
259. v6_conmetr1(sK0) <- (3, 4, 5, ~6, 17) [subsumption resolution 258,131] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
260. 18 | ~3 | ~4 | ~5 | 6 | ~17 [avatar split clause 259,208,129,124,119,114,224]
261. v5_conmetr1(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 16) [subsumption resolution 254,116] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
262. v5_conmetr1(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 4, 16) [subsumption resolution 261,121] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
263. v5_conmetr1(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 16) [subsumption resolution 262,126] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
264. v5_conmetr1(sK0) <- (3, 4, 5, ~6, 16) [subsumption resolution 263,131] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
265. 15 | ~3 | ~4 | ~5 | 6 | ~16 [avatar split clause 264,199,129,124,119,114,188]
272. v1_aff_2(sK0) | ~v2_aff_2(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 16) [subsumption resolution 253,116] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
273. v1_aff_2(sK0) | ~v2_aff_2(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 4, 16) [subsumption resolution 272,121] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
274. v1_aff_2(sK0) | ~v2_aff_2(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 16) [subsumption resolution 273,126] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
275. v1_aff_2(sK0) | ~v2_aff_2(sK0) <- (3, 4, 5, ~6, 16) [subsumption resolution 274,131] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
276. ~17 | 1 | ~3 | ~4 | ~5 | 6 | ~16 [avatar split clause 275,199,129,124,119,114,104,208]
277. v6_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2) [resolution 109,99] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
278. v5_conmetr1(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2) [resolution 109,98] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
279. v6_conmetr1(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3) [subsumption resolution 277,116] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
280. v6_conmetr1(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3, 4) [subsumption resolution 279,121] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
281. v6_conmetr1(sK0) | v7_struct_0(sK0) <- (2, 3, 4, 5) [subsumption resolution 280,126] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
282. v6_conmetr1(sK0) <- (2, 3, 4, 5, ~6) [subsumption resolution 281,131] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
283. 18 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 282,129,124,119,114,108,224]
284. v5_conmetr1(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3) [subsumption resolution 278,116] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
285. v5_conmetr1(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 3, 4) [subsumption resolution 284,121] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
286. v5_conmetr1(sK0) | v7_struct_0(sK0) <- (2, 3, 4, 5) [subsumption resolution 285,126] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
287. v5_conmetr1(sK0) <- (2, 3, 4, 5, ~6) [subsumption resolution 286,131] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
288. 15 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 287,129,124,119,114,108,188]
292. v2_aff_2(sK0) | ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (18) [resolution 226,94] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
294. ~l1_analoaf(sK0) | ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (~17, 18) [subsumption resolution 292,209] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
295. ~v2_diraf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, ~17, 18) [subsumption resolution 294,116] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
296. ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 4, ~17, 18) [subsumption resolution 295,121] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
297. v7_struct_0(sK0) <- (3, 4, 5, ~17, 18) [subsumption resolution 296,126] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
298. $false <- (3, 4, 5, ~6, ~17, 18) [subsumption resolution 297,131] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
299. ~3 | ~4 | ~5 | 6 | 17 | ~18 [avatar contradiction clause 298]
300. $false [avatar sat refutation 111,112,117,122,127,132,234,239,245,252,260,265,276,283,288,299] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:212,thDist:-212}
% SZS output end Proof for conmetr1__t11_conmetr1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 33
% Generated clauses: 160
% Active clauses: 44
% Passive clauses: 51
% Final active clauses: 35
% Main loop iterations started: 48
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 67
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 25
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 31
% SAT solver unit clauses: 6
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

