i: [1,39,0,2,1,1,0,0] t153_relat_1
a: [3,73,0,2,1,39]
s: 73
i: [1,38,0,2,1,1,0,0] t153_relat_1
a: [3,78,0,2,1,38]
s: 78
i: [1,40,0,3,1,1,0,0] t153_relat_1
a: [3,68,0,3,1,40]
s: 68
i: [1,42,0,7,1,1,0,0] t153_relat_1
a: [3,58,0,7,1,42]
s: 58
i: [1,41,0,7,1,1,0,0] t153_relat_1
a: [3,63,0,7,1,41]
s: 63
i: [1,49,0,3,1,0,0,1] reflexivity_r1_tarski
s: 49
i: [1,48,0,4,1,0,0,255] existence_m1_subset_1
s: 48
i: [1,50,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 50
i: [1,52,0,6,2,0,0,2] fc1_relat_1
s: 52
i: [1,53,0,6,2,0,0,1] dt_k5_relat_1
s: 53
i: [1,45,0,7,2,0,0,255] t3_subset
s: 45
d: [2,81,1,6,2,1,44,53,63]
d: [2,82,1,4,1,2,21,81,73]
a: [3,86,1,4,1,82]
s: 86
i: [1,46,0,7,2,0,0,255] t3_subset
s: 46
d: [2,88,1,5,1,0,40,45,48]
s: 88
i: [1,51,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 51
i: [1,43,0,8,3,0,0,255] cc2_relat_1
s: 43
d: [2,90,1,6,2,0,44,52,51]
s: 90
i: [1,44,0,8,2,0,0,1] t28_xboole_1
s: 44
d: [2,91,1,6,2,0,44,52,51]
s: 91
i: [1,54,0,11,2,0,0,1] rd5_relat_1
s: 54
i: [1,47,0,13,2,0,0,1] t71_relat_1
s: 47
d: [2,103,1,11,2,1,44,54,63]
d: [2,106,1,9,1,2,21,103,73]
a: [3,110,1,9,1,106]
s: 110
d: [2,112,1,13,2,1,44,47,63]
d: [2,121,1,11,1,2,21,112,73]
s: 121
d: [2,141,2,10,2,2,44,53,121]
d: [2,153,2,6,1,3,21,141,86]
s: 153
d: [2,92,1,6,2,0,40,43,48]
s: 92
d: [2,93,1,7,3,0,40,43,46]
s: 93
d: [2,160,2,4,2,1,40,93,68]
a: [3,170,2,2,1,160]
s: 170
d: [2,97,1,8,2,0,44,44,51]
s: 97
d: [2,156,3,6,1,3,44,153,51]
s: 156
d: [2,98,1,8,2,0,44,44,51]
s: 98
d: [2,100,1,8,2,0,44,51,44]
s: 100
d: [2,157,3,6,1,3,44,153,51]
s: 157
d: [2,101,1,8,2,0,44,51,44]
s: 101
d: [2,180,3,7,2,3,44,153,97]
s: 180
d: [2,113,1,19,3,0,44,47,54]
d: [2,122,1,15,2,0,21,113,53]
s: 122
d: [2,117,1,19,3,0,44,54,47]
d: [2,124,1,15,2,0,21,117,53]
s: 124
d: [2,183,4,7,2,3,44,156,44]
s: 183
d: [2,114,1,21,3,0,44,47,47]
d: [2,123,1,17,2,0,21,114,53]
s: 123
d: [2,135,2,13,2,2,44,121,47]
d: [2,145,2,11,1,3,21,135,78]
s: 145
r: 153
r: 121
d: [2,292,1,11,1,3,23,121,145]
s: 292
d: [2,291,2,6,1,4,23,153,145]
s: 291
d: [2,323,2,10,2,3,44,53,292]
d: [2,336,2,6,1,4,21,323,86]
s: 336
d: [2,138,2,13,2,2,44,47,121]
d: [2,149,2,11,1,3,21,138,78]
s: 149
d: [2,305,3,8,2,3,44,53,145]
d: [2,313,3,6,1,3,21,305,73]
s: 313
d: [2,127,2,17,2,1,44,47,110]
d: [2,132,2,13,1,2,21,127,86]
s: 132
d: [2,140,2,17,2,2,44,54,121]
d: [2,152,2,13,1,3,21,140,86]
s: 152
d: [2,343,3,6,1,4,44,291,51]
s: 343
d: [2,189,2,15,2,1,44,122,63]
d: [2,207,2,13,1,2,21,189,73]
d: [2,208,2,13,1,2,22,207,121]
s: 208
d: [2,344,3,6,1,4,44,291,51]
s: 344
d: [2,193,2,17,2,1,44,122,110]
d: [2,211,2,13,1,2,21,193,86]
d: [2,212,2,13,1,4,22,211,121]
s: 212
d: [2,199,2,15,2,2,44,122,121]
d: [2,214,2,13,1,3,21,199,78]
d: [2,215,2,13,1,3,22,214,121]
s: 215
d: [2,341,3,7,2,4,44,291,97]
s: 341
d: [2,201,2,15,2,2,44,121,122]
d: [2,216,2,13,1,3,21,201,78]
d: [2,217,2,13,1,3,22,216,121]
s: 217
d: [2,349,4,7,2,4,44,343,44]
s: 349
d: [2,316,2,15,2,3,44,292,122]
d: [2,327,2,13,1,3,21,316,78]
d: [2,328,2,13,1,3,22,327,292]
s: 328
d: [2,319,2,15,2,3,44,122,292]
d: [2,330,2,13,1,3,21,319,78]
d: [2,331,2,13,1,3,22,330,292]
s: 331
d: [2,294,3,10,2,3,44,145,97]
s: 294
d: [2,374,4,0,0,5,52,68,58,294]
e: 374
d: [2,377,4,10,2,4,44,58,294]
d: [2,384,4,3,1,4,19,377]
d: [2,387,4,0,0,5,21,384,68]
e: 387
f: 387,38,39,40,41,42
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t153_relat_1
% SZS output start Proof for relat_1__t153_relat_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2] : ! [X3] : ((k5_relat_1(X0,X3) = k5_relat_1(X1,X3) & r1_tarski(X2,X3)) => k5_relat_1(X0,X2) = k5_relat_1(X1,X2)))) [input t153_relat_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2] : ! [X3] : ((k5_relat_1(X0,X3) = k5_relat_1(X1,X3) & r1_tarski(X2,X3)) => k5_relat_1(X0,X2) = k5_relat_1(X1,X2)))) [negated conjecture 1]
4. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
14. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
16. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1))) [input t71_relat_1]
17. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2,X3] : ((k5_relat_1(X0,X3) = k5_relat_1(X1,X3) & r1_tarski(X2,X3)) => k5_relat_1(X0,X2) = k5_relat_1(X1,X2)))) [flattening 2]
18. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 14]
20. ! [X0,X1,X2] : (v1_relat_1(X2) => k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1))) [flattening 16]
23. ? [X0] : (? [X1] : (? [X2,X3] : (k5_relat_1(X0,X2) != k5_relat_1(X1,X2) & (k5_relat_1(X0,X3) = k5_relat_1(X1,X3) & r1_tarski(X2,X3))) & v1_relat_1(X1)) & v1_relat_1(X0)) [ennf transformation 17]
24. ? [X0] : (? [X1] : (? [X2,X3] : (k5_relat_1(X0,X2) != k5_relat_1(X1,X2) & k5_relat_1(X0,X3) = k5_relat_1(X1,X3) & r1_tarski(X2,X3)) & v1_relat_1(X1)) & v1_relat_1(X0)) [flattening 23]
26. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 18]
27. ! [X0,X1,X2] : (k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1)) | ~v1_relat_1(X2)) [ennf transformation 20]
31. ? [X0] : (? [X1] : (? [X2,X3] : (k5_relat_1(X0,X2) != k5_relat_1(X1,X2) & k5_relat_1(X0,X3) = k5_relat_1(X1,X3) & r1_tarski(X2,X3)) & v1_relat_1(X1)) & v1_relat_1(X0)) => (? [X1] : (? [X3,X2] : (k5_relat_1(X1,X2) != k5_relat_1(sK0,X2) & k5_relat_1(X1,X3) = k5_relat_1(sK0,X3) & r1_tarski(X2,X3)) & v1_relat_1(X1)) & v1_relat_1(sK0)) [choice axiom]
32. ? [X1] : (? [X3,X2] : (k5_relat_1(X1,X2) != k5_relat_1(sK0,X2) & k5_relat_1(X1,X3) = k5_relat_1(sK0,X3) & r1_tarski(X2,X3)) & v1_relat_1(X1)) => (? [X3,X2] : (k5_relat_1(sK0,X2) != k5_relat_1(sK1,X2) & k5_relat_1(sK0,X3) = k5_relat_1(sK1,X3) & r1_tarski(X2,X3)) & v1_relat_1(sK1)) [choice axiom]
33. ? [X3,X2] : (k5_relat_1(sK0,X2) != k5_relat_1(sK1,X2) & k5_relat_1(sK0,X3) = k5_relat_1(sK1,X3) & r1_tarski(X2,X3)) => (k5_relat_1(sK0,sK2) != k5_relat_1(sK1,sK2) & k5_relat_1(sK0,sK3) = k5_relat_1(sK1,sK3) & r1_tarski(sK2,sK3)) [choice axiom]
34. ((k5_relat_1(sK0,sK2) != k5_relat_1(sK1,sK2) & k5_relat_1(sK0,sK3) = k5_relat_1(sK1,sK3) & r1_tarski(sK2,sK3)) & v1_relat_1(sK1)) & v1_relat_1(sK0) [skolemisation 24,33,32,31]
38. v1_relat_1(sK0) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
39. v1_relat_1(sK1) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
40. r1_tarski(sK2,sK3) [cnf transformation 34] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
41. k5_relat_1(sK0,sK3) = k5_relat_1(sK1,sK3) [cnf transformation 34] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
42. k5_relat_1(sK0,sK2) != k5_relat_1(sK1,sK2) [cnf transformation 34] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
44. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. k5_relat_1(k5_relat_1(X2,X0),X1) = k5_relat_1(X2,k3_xboole_0(X0,X1)) | ~v1_relat_1(X2) [cnf transformation 27] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
56. 1 <=> k5_relat_1(sK0,sK2) = k5_relat_1(sK1,sK2) [avatar definition]
58. k5_relat_1(sK0,sK2) != k5_relat_1(sK1,sK2) <- (~1) [avatar component clause 56] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
59. ~1 [avatar split clause 42,56]
61. 2 <=> k5_relat_1(sK0,sK3) = k5_relat_1(sK1,sK3) [avatar definition]
63. k5_relat_1(sK0,sK3) = k5_relat_1(sK1,sK3) <- (2) [avatar component clause 61] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
64. 2 [avatar split clause 41,61]
66. 3 <=> r1_tarski(sK2,sK3) [avatar definition]
68. r1_tarski(sK2,sK3) <- (3) [avatar component clause 66] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
69. 3 [avatar split clause 40,66]
71. 4 <=> v1_relat_1(sK1) [avatar definition]
73. v1_relat_1(sK1) <- (4) [avatar component clause 71] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
74. 4 [avatar split clause 39,71]
76. 5 <=> v1_relat_1(sK0) [avatar definition]
78. v1_relat_1(sK0) <- (5) [avatar component clause 76] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
79. 5 [avatar split clause 38,76]
97. k3_xboole_0(X4,X3) = X3 | ~r1_tarski(X3,X4) [superposition 44,51] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
112. k5_relat_1(k5_relat_1(sK0,sK3),X0) = k5_relat_1(sK1,k3_xboole_0(sK3,X0)) | ~v1_relat_1(sK1) <- (2) [superposition 47,63] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
121. k5_relat_1(k5_relat_1(sK0,sK3),X0) = k5_relat_1(sK1,k3_xboole_0(sK3,X0)) <- (2, 4) [subsumption resolution 112,73] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
135. k5_relat_1(sK0,k3_xboole_0(sK3,X1)) = k5_relat_1(sK1,k3_xboole_0(sK3,X1)) | ~v1_relat_1(sK0) <- (2, 4) [superposition 121,47] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
145. k5_relat_1(sK0,k3_xboole_0(sK3,X1)) = k5_relat_1(sK1,k3_xboole_0(sK3,X1)) <- (2, 4, 5) [subsumption resolution 135,78] {a:2,w:11,wCS:1100,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
294. k5_relat_1(sK1,X0) = k5_relat_1(sK0,X0) | ~r1_tarski(X0,sK3) <- (2, 4, 5) [superposition 145,97] {a:3,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
377. k5_relat_1(sK0,sK2) != k5_relat_1(sK0,sK2) | ~r1_tarski(sK2,sK3) <- (~1, 2, 4, 5) [superposition 58,294] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
384. ~r1_tarski(sK2,sK3) <- (~1, 2, 4, 5) [trivial inequality removal 377] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
387. $false <- (~1, 2, 3, 4, 5) [subsumption resolution 384,68] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
388. 1 | ~2 | ~3 | ~4 | ~5 [avatar contradiction clause 387]
394. $false [avatar sat refutation 59,64,69,74,79,388] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:78,thDist:-78}
% SZS output end Proof for relat_1__t153_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 318
% Active clauses: 39
% Passive clauses: 129
% Final active clauses: 37
% Final passive clauses: 72
% Main loop iterations started: 57
% 
% >>> Simplifying Inferences
% Duplicate literals: 15
% Trivial inequalities: 1
% Fw subsumption resolutions: 67
% Fw demodulations: 30
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 60
% Fw demodulations to eq. taut.: 10
% 
% >>> Generating Inferences
% Binary resolution: 7
% Unit resulting resolution: 1
% Forward superposition: 86
% Backward superposition: 79
% Self superposition: 5
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

