i: [1,40,0,8,3,1,0,0] t6_gate_1
a: [3,59,0,2,1,40]
s: 59
a: [3,62,0,2,1,40]
s: 62
a: [3,66,0,4,1,40]
s: 66
i: [1,46,0,2,1,0,0,1] rc2_xboole_0
a: [3,75,0,2,1,46]
s: 75
i: [1,47,0,2,1,0,0,1] rc1_xboole_0
a: [3,80,0,2,1,47]
s: 80
i: [1,41,0,7,1,0,0,255] commutativity_k4_gate_1
s: 41
i: [1,42,0,7,1,0,0,255] commutativity_k4_gate_1
s: 42
i: [1,43,0,8,3,0,0,1] fc9_gate_1
s: 43
d: [2,82,1,0,0,3,52,59,62,66,43]
e: 82
d: [2,83,1,4,2,1,40,43,66]
d: [2,87,1,2,1,2,21,83,62]
d: [2,88,1,0,0,3,21,87,59]
e: 88
r: 59
r: 66
a: [3,58,0,2,1,40]
s: 58
a: [3,67,0,4,1,40]
s: 67
i: [1,44,0,8,3,0,0,1] fc7_gate_1
s: 44
d: [2,90,1,0,0,3,52,58,62,67,44]
e: 90
d: [2,92,1,4,2,1,40,44,67]
d: [2,97,1,2,1,2,21,92,62]
d: [2,98,1,0,0,3,21,97,58]
e: 98
r: 58
r: 62
s: 59
a: [3,63,0,2,1,40]
s: 63
i: [1,45,0,8,3,0,0,1] fc8_gate_1
s: 45
d: [2,93,1,8,3,0,44,44,41]
s: 93
d: [2,105,2,0,0,3,52,63,59,67,93]
e: 105
d: [2,108,2,4,2,1,40,93,67]
d: [2,114,2,2,1,2,21,108,59]
d: [2,115,2,0,0,3,21,114,63]
e: 115
r: 59
r: 67
s: 58
s: 66
d: [2,117,1,0,0,3,52,58,63,66,45]
e: 117
d: [2,118,1,4,2,1,40,66,45]
d: [2,121,1,2,1,2,21,118,63]
d: [2,122,1,0,0,3,21,121,58]
e: 122
i: [1,39,0,8,3,1,0,0] t6_gate_1
i: [1,32,0,8,3,1,0,0] t6_gate_1
i: [1,29,0,8,3,1,0,0] t6_gate_1
f: 122,115,92,88,29,32,39,40
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t6_gate_1
% SZS output start Proof for gate_1__t6_gate_1
1. ! [X0] : ! [X1] : (~(v1_xboole_0(k4_gate_1(X0,X1)) & ((~v1_xboole_0(X1) & v1_xboole_0(X0)) | (v1_xboole_0(X1) & ~v1_xboole_0(X0)))) & ~(~(~v1_xboole_0(X1) & v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~v1_xboole_0(k4_gate_1(X0,X1)))) [input t6_gate_1]
2. ~! [X0] : ! [X1] : (~(v1_xboole_0(k4_gate_1(X0,X1)) & ((~v1_xboole_0(X1) & v1_xboole_0(X0)) | (v1_xboole_0(X1) & ~v1_xboole_0(X0)))) & ~(~(~v1_xboole_0(X1) & v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~v1_xboole_0(k4_gate_1(X0,X1)))) [negated conjecture 1]
3. ! [X0,X1] : k4_gate_1(X0,X1) = k4_gate_1(X1,X0) [input commutativity_k4_gate_1]
7. ! [X0,X1] : ((~v1_xboole_0(X1) & v1_xboole_0(X0)) => ~v1_xboole_0(k4_gate_1(X0,X1))) [input fc7_gate_1]
8. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k4_gate_1(X0,X1))) [input fc8_gate_1]
9. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => v1_xboole_0(k4_gate_1(X0,X1))) [input fc9_gate_1]
12. ~! [X0,X1] : (~(v1_xboole_0(k4_gate_1(X0,X1)) & ((~v1_xboole_0(X1) & v1_xboole_0(X0)) | (v1_xboole_0(X1) & ~v1_xboole_0(X0)))) & ~(~(~v1_xboole_0(X1) & v1_xboole_0(X0)) & ~(v1_xboole_0(X1) & ~v1_xboole_0(X0)) & ~v1_xboole_0(k4_gate_1(X0,X1)))) [flattening 2]
13. ? [X0,X1] : ((v1_xboole_0(k4_gate_1(X0,X1)) & ((~v1_xboole_0(X1) & v1_xboole_0(X0)) | (v1_xboole_0(X1) & ~v1_xboole_0(X0)))) | ((v1_xboole_0(X1) | ~v1_xboole_0(X0)) & (~v1_xboole_0(X1) | v1_xboole_0(X0)) & ~v1_xboole_0(k4_gate_1(X0,X1)))) [ennf transformation 12]
14. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 9]
15. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 14]
16. ! [X0,X1] : (~v1_xboole_0(k4_gate_1(X0,X1)) | (v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 7]
17. ! [X0,X1] : (~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 16]
18. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 8]
19. ! [X0,X1] : (v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 18]
20. ? [X0,X1] : ((v1_xboole_0(k4_gate_1(X0,X1)) & ((~v1_xboole_0(X1) & v1_xboole_0(X0)) | (v1_xboole_0(X1) & ~v1_xboole_0(X0)))) | ((v1_xboole_0(X1) | ~v1_xboole_0(X0)) & (~v1_xboole_0(X1) | v1_xboole_0(X0)) & ~v1_xboole_0(k4_gate_1(X0,X1)))) => ((v1_xboole_0(k4_gate_1(sK0,sK1)) & ((~v1_xboole_0(sK1) & v1_xboole_0(sK0)) | (v1_xboole_0(sK1) & ~v1_xboole_0(sK0)))) | ((v1_xboole_0(sK1) | ~v1_xboole_0(sK0)) & (~v1_xboole_0(sK1) | v1_xboole_0(sK0)) & ~v1_xboole_0(k4_gate_1(sK0,sK1)))) [choice axiom]
21. (v1_xboole_0(k4_gate_1(sK0,sK1)) & ((~v1_xboole_0(sK1) & v1_xboole_0(sK0)) | (v1_xboole_0(sK1) & ~v1_xboole_0(sK0)))) | ((v1_xboole_0(sK1) | ~v1_xboole_0(sK0)) & (~v1_xboole_0(sK1) | v1_xboole_0(sK0)) & ~v1_xboole_0(k4_gate_1(sK0,sK1))) [skolemisation 13,20]
29. v1_xboole_0(sK0) | v1_xboole_0(sK1) | ~v1_xboole_0(k4_gate_1(sK0,sK1)) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) | ~v1_xboole_0(k4_gate_1(sK0,sK1)) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. v1_xboole_0(k4_gate_1(sK0,sK1)) | ~v1_xboole_0(sK1) | v1_xboole_0(sK0) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. v1_xboole_0(k4_gate_1(sK0,sK1)) | v1_xboole_0(sK1) | ~v1_xboole_0(sK0) [cnf transformation 21] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
41. k4_gate_1(X0,X1) = k4_gate_1(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
43. v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 15] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. ~v1_xboole_0(k4_gate_1(X0,X1)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 17] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. v1_xboole_0(k4_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. 1 <=> v1_xboole_0(sK0) [avatar definition]
58. v1_xboole_0(sK0) <- (1) [avatar component clause 57] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
59. ~v1_xboole_0(sK0) <- (~1) [avatar component clause 57] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
61. 2 <=> v1_xboole_0(sK1) [avatar definition]
62. ~v1_xboole_0(sK1) <- (~2) [avatar component clause 61] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. v1_xboole_0(sK1) <- (2) [avatar component clause 61] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. 3 <=> v1_xboole_0(k4_gate_1(sK0,sK1)) [avatar definition]
66. ~v1_xboole_0(k4_gate_1(sK0,sK1)) <- (~3) [avatar component clause 65] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
67. v1_xboole_0(k4_gate_1(sK0,sK1)) <- (3) [avatar component clause 65] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. ~1 | 2 | 3 [avatar split clause 40,65,61,57]
69. 1 | ~2 | 3 [avatar split clause 39,65,61,57]
70. ~3 | ~1 | ~2 [avatar split clause 32,61,57,65]
71. ~3 | 2 | 1 [avatar split clause 29,57,61,65]
83. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~3) [resolution 43,66] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
87. v1_xboole_0(sK0) <- (~2, ~3) [subsumption resolution 83,62] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
88. $false <- (~1, ~2, ~3) [subsumption resolution 87,59] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
89. 1 | 2 | 3 [avatar contradiction clause 88]
92. v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (3) [resolution 44,67] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
93. ~v1_xboole_0(k4_gate_1(X1,X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [superposition 44,41] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
100. ~1 | 2 | ~3 [avatar split clause 92,65,61,57]
108. v1_xboole_0(sK0) | ~v1_xboole_0(sK1) <- (3) [resolution 93,67] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
114. ~v1_xboole_0(sK1) <- (~1, 3) [subsumption resolution 108,59] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
115. $false <- (~1, 2, 3) [subsumption resolution 114,63] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
116. 1 | ~2 | ~3 [avatar contradiction clause 115]
118. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~3) [resolution 66,45] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
121. ~v1_xboole_0(sK0) <- (2, ~3) [subsumption resolution 118,63] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
122. $false <- (1, 2, ~3) [subsumption resolution 121,58] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
123. ~1 | ~2 | 3 [avatar contradiction clause 122]
124. $false [avatar sat refutation 68,69,70,71,89,100,116,123] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:74,thDist:-74}
% SZS output end Proof for gate_1__t6_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 22
% Generated clauses: 77
% Active clauses: 16
% Passive clauses: 18
% Final active clauses: 10
% Final passive clauses: 1
% Main loop iterations started: 17
% 
% >>> Simplifying Inferences
% Duplicate literals: 12
% Fw subsumption resolutions: 8
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 4
% Forward superposition: 8
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 14
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 2
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

