i: [1,269,0,2,1,1,0,0] t22_connsp_2
a: [3,400,0,2,1,269]
s: 400
i: [1,268,0,2,1,1,0,0] t22_connsp_2
a: [3,405,0,2,1,268]
s: 405
i: [1,267,0,2,1,1,0,0] t22_connsp_2
a: [3,410,0,2,1,267]
s: 410
i: [1,276,0,3,1,1,0,0] t22_connsp_2
a: [3,385,0,3,1,276]
s: 385
i: [1,270,0,4,1,1,0,0] t22_connsp_2
a: [3,395,0,4,1,270]
s: 395
i: [1,271,0,6,1,1,0,0] t22_connsp_2
a: [3,390,0,6,1,271]
s: 390
d: [2,442,0,6,1,2,22,390,385]
a: [3,446,0,6,1,442]
s: 446
i: [1,277,0,2,1,0,0,255] spc1_boole
a: [3,415,0,2,1,277]
s: 415
i: [1,278,0,2,1,0,0,1] fc1_xboole_0
a: [3,420,0,2,1,278]
s: 420
i: [1,378,0,2,1,0,0,2] rc2_xboole_0
a: [3,425,0,2,1,378]
s: 425
i: [1,379,0,2,1,0,0,2] existence_l1_struct_0
a: [3,430,0,2,1,379]
s: 430
i: [1,380,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,435,0,2,1,380]
s: 435
i: [1,381,0,2,1,0,0,2] rc1_xboole_0
a: [3,440,0,2,1,381]
s: 440
i: [1,272,0,11,3,1,0,0] t22_connsp_2
s: 272
i: [1,273,0,11,3,1,0,0] t22_connsp_2
s: 273
i: [1,274,0,11,3,1,0,0] t22_connsp_2
s: 274
d: [2,448,0,11,3,1,22,272,385]
s: 448
d: [2,449,0,11,3,1,22,273,385]
s: 449
d: [2,450,0,11,3,1,22,274,385]
s: 450
i: [1,279,0,3,1,0,0,2] fc1_subset_1
s: 279
i: [1,373,0,3,1,0,0,2] rc2_subset_1
s: 373
i: [1,376,0,3,1,0,0,255] reflexivity_r1_tarski
s: 376
i: [1,291,0,4,2,0,0,1] dt_l1_pre_topc
s: 291
d: [2,451,1,2,1,1,40,291,400]
a: [3,456,1,2,1,451]
s: 456
i: [1,371,0,4,1,0,0,1] existence_m1_subset_1
s: 371
i: [1,375,0,4,1,0,0,255] rc3_subset_1
s: 375
d: [2,463,1,12,2,1,40,371,450]
a: [3,469,1,6,1,463]
s: 469
i: [1,275,0,17,5,1,0,0] t22_connsp_2
s: 275
d: [2,452,1,2,1,1,40,291,435]
a: [3,461,1,2,1,452]
s: 461
d: [2,485,0,17,5,1,22,275,385]
s: 485
i: [1,281,0,5,2,0,0,255] rc6_subset_1
s: 281
a: [3,472,1,6,1,463]
s: 472
i: [1,285,0,5,2,0,0,2] rc1_subset_1
s: 285
i: [1,287,0,5,2,0,0,255] rc5_subset_1
s: 287
i: [1,288,0,5,2,0,0,255] rc5_subset_1
s: 288
i: [1,311,0,5,2,0,0,1] rc1_connsp_1
s: 311
i: [1,316,0,5,2,0,0,1] t6_boole
s: 316
i: [1,368,0,5,2,0,0,1] t7_boole
s: 368
d: [2,492,1,8,2,1,44,446,316]
s: 492
i: [1,372,0,5,1,0,0,2] rc2_subset_1
s: 372
i: [1,374,0,5,1,0,0,255] rc3_subset_1
s: 374
d: [2,493,1,20,4,1,40,372,485]
a: [3,499,1,5,1,493]
s: 499
i: [1,290,0,6,2,0,0,255] rc4_subset_1
s: 290
a: [3,512,1,5,1,493]
s: 512
i: [1,295,0,6,3,0,0,255] cc7_struct_0
s: 295
i: [1,296,0,6,3,0,0,255] cc5_struct_0
s: 296
d: [2,517,1,20,4,1,40,374,485]
a: [3,523,1,5,1,517]
s: 523
i: [1,298,0,6,3,0,0,255] cc2_struct_0
s: 298
a: [3,536,1,5,1,517]
s: 536
i: [1,299,0,6,3,0,0,255] cc17_struct_0
s: 299
i: [1,300,0,6,3,0,0,255] cc6_struct_0
s: 300
d: [2,542,1,4,2,1,40,296,430]
a: [3,548,1,2,1,542]
s: 548
i: [1,301,0,6,3,0,0,255] cc1_struct_0
s: 301
d: [2,490,2,8,2,1,44,469,316]
s: 490
i: [1,304,0,6,3,0,0,255] cc4_struct_0
s: 304
i: [1,314,0,6,2,0,0,255] rc5_tops_1
s: 314
d: [2,544,2,4,2,1,40,296,461]
a: [3,557,2,2,1,544]
s: 557
i: [1,364,0,6,2,0,0,1] t1_subset
s: 364
i: [1,377,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 377
i: [1,280,0,7,2,0,0,255] rc6_subset_1
s: 280
i: [1,284,0,7,2,0,0,2] rc1_subset_1
s: 284
d: [2,577,1,23,5,1,40,280,485]
a: [3,583,1,5,1,577]
s: 583
i: [1,286,0,7,2,0,0,255] rc5_subset_1
s: 286
a: [3,596,1,5,1,577]
s: 596
i: [1,289,0,7,2,0,0,255] rc4_subset_1
s: 289
i: [1,292,0,7,2,0,0,255] fc12_struct_0
s: 292
d: [2,605,1,23,5,1,40,284,485]
a: [3,611,1,5,1,605]
s: 611
i: [1,293,0,7,2,0,0,1] d3_struct_0
s: 293
a: [3,624,1,5,1,605]
s: 624
i: [1,302,0,7,3,0,0,255] cc8_struct_0
s: 302
i: [1,305,0,7,3,0,0,255] cc11_struct_0
s: 305
d: [2,634,1,23,5,1,40,286,485]
a: [3,640,1,5,1,634]
s: 640
i: [1,306,0,7,3,0,0,255] cc11_struct_0
s: 306
a: [3,653,1,5,1,634]
s: 653
i: [1,307,0,7,3,0,0,255] cc9_struct_0
s: 307
i: [1,319,0,7,3,0,0,255] fc9_struct_0
s: 319
d: [2,658,1,23,5,1,40,289,485]
a: [3,664,1,5,1,658]
s: 664
i: [1,320,0,7,3,0,0,255] fc6_struct_0
s: 320
a: [3,677,1,5,1,658]
s: 677
i: [1,322,0,7,3,0,0,255] rc22_struct_0
s: 322
i: [1,334,0,7,3,0,0,2] fc5_struct_0
s: 334
i: [1,335,0,7,3,0,0,1] fc2_struct_0
s: 335
i: [1,337,0,7,3,0,0,1] rc4_struct_0
s: 337
i: [1,339,0,7,3,0,0,255] rc21_struct_0
s: 339
i: [1,340,0,7,3,0,0,255] rc21_struct_0
s: 340
i: [1,360,0,7,3,0,0,255] fc8_struct_0
s: 360
i: [1,361,0,7,3,0,0,255] fc7_struct_0
s: 361
i: [1,362,0,7,3,0,0,2] fc4_struct_0
s: 362
i: [1,363,0,7,3,0,0,1] fc1_struct_0
s: 363
i: [1,366,0,7,2,0,0,255] t3_subset
s: 366
i: [1,367,0,7,3,0,0,1] t8_boole
s: 367
d: [2,704,1,10,3,1,40,366,450]
s: 704
i: [1,294,0,8,2,0,0,2] dt_k2_struct_0
s: 294
d: [2,717,1,18,5,1,40,294,485]
d: [2,723,1,16,4,2,21,717,456]
a: [3,726,1,4,1,723]
s: 726
i: [1,310,0,8,2,0,0,1] rc1_connsp_1
s: 310
i: [1,312,0,8,2,0,0,255] rc4_tops_1
s: 312
a: [3,739,1,4,1,723]
s: 739
i: [1,313,0,8,2,0,0,255] rc5_tops_1
s: 313
d: [2,751,1,18,5,1,40,310,485]
d: [2,755,1,16,4,2,21,751,400]
a: [3,758,1,4,1,755]
s: 758
i: [1,315,0,8,3,0,0,255] cc5_subset_1
s: 315
i: [1,317,0,8,3,0,0,2] cc1_subset_1
s: 317
a: [3,771,1,4,1,755]
s: 771
i: [1,341,0,8,3,0,0,1] fc4_pre_topc
s: 341
d: [2,779,1,18,5,1,40,312,485]
d: [2,783,1,16,4,2,21,779,400]
a: [3,786,1,4,1,783]
s: 786
i: [1,342,0,8,3,0,0,1] fc10_tops_1
s: 342
i: [1,350,0,8,3,0,0,255] rc7_tops_1
s: 350
a: [3,799,1,4,1,783]
s: 799
i: [1,352,0,8,3,0,0,1] rc10_pre_topc
s: 352
d: [2,810,1,18,5,1,40,313,485]
d: [2,814,1,16,4,2,21,810,400]
a: [3,817,1,4,1,814]
s: 817
i: [1,354,0,8,3,0,0,1] rc1_tops_1
s: 354
i: [1,356,0,8,3,0,0,1] rc6_pre_topc
s: 356
a: [3,830,1,4,1,814]
s: 830
i: [1,358,0,8,3,0,0,1] rc2_tops_1
s: 358
d: [2,707,2,8,2,1,40,704,376]
a: [3,711,2,4,1,707]
s: 711
i: [1,359,0,8,3,0,0,1] rc2_tops_1
s: 359
i: [1,365,0,8,3,0,0,1] t2_subset
s: 365
a: [3,628,1,3,1,605]
s: 628
d: [2,921,2,0,0,3,52,410,456,628,335]
e: 921
d: [2,922,2,4,2,1,40,628,335]
d: [2,924,2,2,1,2,21,922,456]
d: [2,925,2,0,0,3,21,924,410]
e: 925
r: 628
r: 739
i: [1,308,0,9,4,0,0,255] cc10_struct_0
s: 308
a: [3,627,1,3,1,605]
s: 627
i: [1,318,0,9,3,0,0,255] cc4_subset_1
s: 318
i: [1,325,0,9,4,0,0,255] rc6_tops_1
s: 325
a: [3,731,1,4,1,723]
s: 731
d: [2,946,2,0,0,3,52,405,400,731,342]
e: 946
d: [2,947,2,4,2,1,40,731,342]
d: [2,950,2,2,1,2,21,947,400]
d: [2,951,2,0,0,3,21,950,405]
e: 951
r: 731
i: [1,328,0,9,4,0,0,1] rc7_pre_topc
s: 328
a: [3,730,1,4,1,723]
s: 730
i: [1,331,0,9,4,0,0,1] rc3_tops_1
s: 331
i: [1,370,0,9,3,0,0,1] t5_subset
s: 370
a: [3,735,1,4,1,723]
s: 735
d: [2,975,2,0,0,3,52,405,400,735,341]
e: 975
d: [2,976,2,4,2,1,40,735,341]
d: [2,979,2,2,1,2,21,976,400]
d: [2,980,2,0,0,3,21,979,405]
e: 980
r: 726
r: 735
i: [1,321,0,10,3,0,0,255] rc22_struct_0
s: 321
a: [3,727,1,4,1,723]
s: 727
d: [2,1021,2,0,0,2,52,420,727,368]
e: 1021
d: [2,1024,2,2,1,1,40,727,368]
d: [2,1028,2,0,0,2,21,1024,420]
e: 1028
d: [2,1025,2,6,2,1,44,727,293]
d: [2,1030,2,2,1,2,21,1025,711]
d: [2,1031,2,0,0,3,21,1030,456]
e: 1031
d: [2,807,2,6,2,1,44,739,293]
d: [2,808,2,4,1,2,21,807,456]
d: [2,891,1,7,2,1,40,365,395]
f: 1028,980,951,925,891,808,723,451,278,267,268,269,270,276
% Refutation found. Thanks to Tanya!
% SZS status Theorem for connsp_2__t22_connsp_2
% SZS output start Proof for connsp_2__t22_connsp_2
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X0)))) => ~(k1_xboole_0 = X2 & ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0))) => (r2_hidden(X3,X2) <=> (r2_hidden(X1,X3) & v4_pre_topc(X3,X0) & v3_pre_topc(X3,X0)))))))) [input t22_connsp_2]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X0)))) => ~(k1_xboole_0 = X2 & ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0))) => (r2_hidden(X3,X2) <=> (r2_hidden(X1,X3) & v4_pre_topc(X3,X0) & v3_pre_topc(X3,X0)))))))) [negated conjecture 1]
29. ! [X0] : (l1_struct_0(X0) => u1_struct_0(X0) = k2_struct_0(X0)) [input d3_struct_0]
32. ! [X0] : (l1_struct_0(X0) => m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0)))) [input dt_k2_struct_0]
33. ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)) [input dt_l1_pre_topc]
40. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => v3_pre_topc(k2_struct_0(X0),X0)) [input fc10_tops_1]
46. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
47. ! [X0] : ((l1_struct_0(X0) & ~v2_struct_0(X0)) => ~v1_xboole_0(u1_struct_0(X0))) [input fc2_struct_0]
48. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => v4_pre_topc(k2_struct_0(X0),X0)) [input fc4_pre_topc]
80. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
85. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
88. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 80]
91. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 85]
100. ? [X0] : (? [X1] : (? [X2] : ((k1_xboole_0 = X2 & ! [X3] : ((r2_hidden(X3,X2) <=> (r2_hidden(X1,X3) & v4_pre_topc(X3,X0) & v3_pre_topc(X3,X0))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0))))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X0))))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
101. ? [X0] : (? [X1] : (? [X2] : (k1_xboole_0 = X2 & ! [X3] : ((r2_hidden(X3,X2) <=> (r2_hidden(X1,X3) & v4_pre_topc(X3,X0) & v3_pre_topc(X3,X0))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0)))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X0))))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [flattening 100]
110. ! [X0] : (l1_struct_0(X0) | ~l1_pre_topc(X0)) [ennf transformation 33]
112. ! [X0] : (u1_struct_0(X0) = k2_struct_0(X0) | ~l1_struct_0(X0)) [ennf transformation 29]
113. ! [X0] : (m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0)) [ennf transformation 32]
161. ! [X0] : (~v1_xboole_0(u1_struct_0(X0)) | (~l1_struct_0(X0) | v2_struct_0(X0))) [ennf transformation 47]
162. ! [X0] : (~v1_xboole_0(u1_struct_0(X0)) | ~l1_struct_0(X0) | v2_struct_0(X0)) [flattening 161]
167. ! [X0] : (v4_pre_topc(k2_struct_0(X0),X0) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 48]
168. ! [X0] : (v4_pre_topc(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 167]
169. ! [X0] : (v3_pre_topc(k2_struct_0(X0),X0) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 40]
170. ! [X0] : (v3_pre_topc(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 169]
202. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 88]
203. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 202]
206. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 91]
211. ? [X0] : (? [X1] : (? [X2] : (k1_xboole_0 = X2 & ! [X3] : (((r2_hidden(X3,X2) | (~r2_hidden(X1,X3) | ~v4_pre_topc(X3,X0) | ~v3_pre_topc(X3,X0))) & ((r2_hidden(X1,X3) & v4_pre_topc(X3,X0) & v3_pre_topc(X3,X0)) | ~r2_hidden(X3,X2))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0)))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X0))))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [nnf transformation 101]
212. ? [X0] : (? [X1] : (? [X2] : (k1_xboole_0 = X2 & ! [X3] : (((r2_hidden(X3,X2) | ~r2_hidden(X1,X3) | ~v4_pre_topc(X3,X0) | ~v3_pre_topc(X3,X0)) & ((r2_hidden(X1,X3) & v4_pre_topc(X3,X0) & v3_pre_topc(X3,X0)) | ~r2_hidden(X3,X2))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0)))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X0))))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [flattening 211]
213. ? [X0] : (? [X1] : (? [X2] : (k1_xboole_0 = X2 & ! [X3] : (((r2_hidden(X3,X2) | ~r2_hidden(X1,X3) | ~v4_pre_topc(X3,X0) | ~v3_pre_topc(X3,X0)) & ((r2_hidden(X1,X3) & v4_pre_topc(X3,X0) & v3_pre_topc(X3,X0)) | ~r2_hidden(X3,X2))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X0)))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X0))))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (k1_xboole_0 = X2 & ! [X3] : (((r2_hidden(X3,X2) | ~r2_hidden(X1,X3) | ~v4_pre_topc(X3,sK0) | ~v3_pre_topc(X3,sK0)) & ((r2_hidden(X1,X3) & v4_pre_topc(X3,sK0) & v3_pre_topc(X3,sK0)) | ~r2_hidden(X3,X2))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(sK0))))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0)) [choice axiom]
214. ? [X1] : (? [X2] : (k1_xboole_0 = X2 & ! [X3] : (((r2_hidden(X3,X2) | ~r2_hidden(X1,X3) | ~v4_pre_topc(X3,sK0) | ~v3_pre_topc(X3,sK0)) & ((r2_hidden(X1,X3) & v4_pre_topc(X3,sK0) & v3_pre_topc(X3,sK0)) | ~r2_hidden(X3,X2))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(sK0))))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (k1_xboole_0 = X2 & ! [X3] : (((r2_hidden(X3,X2) | ~r2_hidden(sK1,X3) | ~v4_pre_topc(X3,sK0) | ~v3_pre_topc(X3,sK0)) & ((r2_hidden(sK1,X3) & v4_pre_topc(X3,sK0) & v3_pre_topc(X3,sK0)) | ~r2_hidden(X3,X2))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(sK0))))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
215. ? [X2] : (k1_xboole_0 = X2 & ! [X3] : (((r2_hidden(X3,X2) | ~r2_hidden(sK1,X3) | ~v4_pre_topc(X3,sK0) | ~v3_pre_topc(X3,sK0)) & ((r2_hidden(sK1,X3) & v4_pre_topc(X3,sK0) & v3_pre_topc(X3,sK0)) | ~r2_hidden(X3,X2))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(sK0))))) => (k1_xboole_0 = sK2 & ! [X3] : (((r2_hidden(X3,sK2) | ~r2_hidden(sK1,X3) | ~v4_pre_topc(X3,sK0) | ~v3_pre_topc(X3,sK0)) & ((r2_hidden(sK1,X3) & v4_pre_topc(X3,sK0) & v3_pre_topc(X3,sK0)) | ~r2_hidden(X3,sK2))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(sK0))))) [choice axiom]
216. ((k1_xboole_0 = sK2 & ! [X3] : (((r2_hidden(X3,sK2) | ~r2_hidden(sK1,X3) | ~v4_pre_topc(X3,sK0) | ~v3_pre_topc(X3,sK0)) & ((r2_hidden(sK1,X3) & v4_pre_topc(X3,sK0) & v3_pre_topc(X3,sK0)) | ~r2_hidden(X3,sK2))) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(sK0)))) & m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(sK0))))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0) [skolemisation 212,215,214,213]
267. ~v2_struct_0(sK0) [cnf transformation 216] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
268. v2_pre_topc(sK0) [cnf transformation 216] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
269. l1_pre_topc(sK0) [cnf transformation 216] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
270. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 216] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
275. r2_hidden(X3,sK2) | ~r2_hidden(sK1,X3) | ~v4_pre_topc(X3,sK0) | ~v3_pre_topc(X3,sK0) | ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 216] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
276. k1_xboole_0 = sK2 [cnf transformation 216] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
278. v1_xboole_0(k1_xboole_0) [cnf transformation 46] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
291. ~l1_pre_topc(X0) | l1_struct_0(X0) [cnf transformation 110] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
293. u1_struct_0(X0) = k2_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 112] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
294. m1_subset_1(k2_struct_0(X0),k1_zfmisc_1(u1_struct_0(X0))) | ~l1_struct_0(X0) [cnf transformation 113] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
335. ~v1_xboole_0(u1_struct_0(X0)) | ~l1_struct_0(X0) | v2_struct_0(X0) [cnf transformation 162] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
341. v4_pre_topc(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 168] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
342. v3_pre_topc(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 170] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
365. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 203] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
368. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 206] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
383. 1 <=> k1_xboole_0 = sK2 [avatar definition]
385. k1_xboole_0 = sK2 <- (1) [avatar component clause 383] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
386. 1 [avatar split clause 276,383]
393. 3 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
395. m1_subset_1(sK1,u1_struct_0(sK0)) <- (3) [avatar component clause 393] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
396. 3 [avatar split clause 270,393]
398. 4 <=> l1_pre_topc(sK0) [avatar definition]
400. l1_pre_topc(sK0) <- (4) [avatar component clause 398] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
401. 4 [avatar split clause 269,398]
403. 5 <=> v2_pre_topc(sK0) [avatar definition]
405. v2_pre_topc(sK0) <- (5) [avatar component clause 403] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
406. 5 [avatar split clause 268,403]
408. 6 <=> v2_struct_0(sK0) [avatar definition]
410. ~v2_struct_0(sK0) <- (~6) [avatar component clause 408] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
411. ~6 [avatar split clause 267,408]
418. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
420. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 418] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
421. 8 [avatar split clause 278,418]
451. l1_struct_0(sK0) <- (4) [resolution 291,400] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
454. 14 <=> l1_struct_0(sK0) [avatar definition]
456. l1_struct_0(sK0) <- (14) [avatar component clause 454] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
457. 14 | ~4 [avatar split clause 451,398,454]
485. ~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(sK0))) | ~r2_hidden(sK1,X3) | ~v4_pre_topc(X3,sK0) | ~v3_pre_topc(X3,sK0) | r2_hidden(X3,k1_xboole_0) <- (1) [forward demodulation 275,385] {a:0,w:17,wCS:1700,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
626. 43 <=> v1_xboole_0(u1_struct_0(sK0)) [avatar definition]
628. v1_xboole_0(u1_struct_0(sK0)) <- (43) [avatar component clause 626] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
713. 53 <=> r2_hidden(sK1,u1_struct_0(sK0)) [avatar definition]
717. ~l1_struct_0(sK0) | ~r2_hidden(sK1,k2_struct_0(sK0)) | ~v4_pre_topc(k2_struct_0(sK0),sK0) | ~v3_pre_topc(k2_struct_0(sK0),sK0) | r2_hidden(k2_struct_0(sK0),k1_xboole_0) <- (1) [resolution 294,485] {a:1,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
723. ~r2_hidden(sK1,k2_struct_0(sK0)) | ~v4_pre_topc(k2_struct_0(sK0),sK0) | ~v3_pre_topc(k2_struct_0(sK0),sK0) | r2_hidden(k2_struct_0(sK0),k1_xboole_0) <- (1, 14) [subsumption resolution 717,456] {a:1,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
725. 54 <=> r2_hidden(k2_struct_0(sK0),k1_xboole_0) [avatar definition]
727. r2_hidden(k2_struct_0(sK0),k1_xboole_0) <- (54) [avatar component clause 725] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
729. 55 <=> v3_pre_topc(k2_struct_0(sK0),sK0) [avatar definition]
731. ~v3_pre_topc(k2_struct_0(sK0),sK0) <- (~55) [avatar component clause 729] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
733. 56 <=> v4_pre_topc(k2_struct_0(sK0),sK0) [avatar definition]
735. ~v4_pre_topc(k2_struct_0(sK0),sK0) <- (~56) [avatar component clause 733] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
737. 57 <=> r2_hidden(sK1,k2_struct_0(sK0)) [avatar definition]
739. ~r2_hidden(sK1,k2_struct_0(sK0)) <- (~57) [avatar component clause 737] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
740. 54 | ~55 | ~56 | ~57 | ~1 | ~14 [avatar split clause 723,454,383,737,733,729,725]
807. ~r2_hidden(sK1,u1_struct_0(sK0)) | ~l1_struct_0(sK0) <- (~57) [superposition 739,293] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
808. ~r2_hidden(sK1,u1_struct_0(sK0)) <- (14, ~57) [subsumption resolution 807,456] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
809. ~53 | ~14 | 57 [avatar split clause 808,737,454,713]
891. v1_xboole_0(u1_struct_0(sK0)) | r2_hidden(sK1,u1_struct_0(sK0)) <- (3) [resolution 365,395] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
911. 53 | 43 | ~3 [avatar split clause 891,393,626,713]
922. ~l1_struct_0(sK0) | v2_struct_0(sK0) <- (43) [resolution 628,335] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
924. v2_struct_0(sK0) <- (14, 43) [subsumption resolution 922,456] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
925. $false <- (~6, 14, 43) [subsumption resolution 924,410] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
926. 6 | ~14 | ~43 [avatar contradiction clause 925]
947. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~55) [resolution 731,342] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
950. ~v2_pre_topc(sK0) <- (4, ~55) [subsumption resolution 947,400] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
951. $false <- (4, 5, ~55) [subsumption resolution 950,405] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
952. ~4 | ~5 | 55 [avatar contradiction clause 951]
976. ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (~56) [resolution 735,341] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
979. ~v2_pre_topc(sK0) <- (4, ~56) [subsumption resolution 976,400] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
980. $false <- (4, 5, ~56) [subsumption resolution 979,405] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
981. ~4 | ~5 | 56 [avatar contradiction clause 980]
1024. ~v1_xboole_0(k1_xboole_0) <- (54) [resolution 727,368] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1028. $false <- (8, 54) [subsumption resolution 1024,420] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1029. ~8 | ~54 [avatar contradiction clause 1028]
1034. $false [avatar sat refutation 386,396,401,406,411,421,457,740,809,911,926,952,981,1029] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:733,thDist:-733}
% SZS output end Proof for connsp_2__t22_connsp_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 85
% >>> Preprocessing
% Introduced skolems: 28
% Pure predicates: 3
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 115
% Generated clauses: 423
% Active clauses: 119
% Passive clauses: 225
% Final active clauses: 114
% Final passive clauses: 95
% Main loop iterations started: 129
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 52
% Fw demodulations: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 21
% Forward subsumptions: 30
% 
% >>> Generating Inferences
% Binary resolution: 151
% Unit resulting resolution: 4
% Forward superposition: 25
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 56
% Split components: 155
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 86
% SAT solver unit clauses: 12
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9978
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% resolution: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

