i: [1,101,0,2,1,1,0,0] t47_subset_1
a: [3,146,0,2,1,101]
s: 146
i: [1,99,0,2,1,1,0,0] t47_subset_1
a: [3,156,0,2,1,99]
s: 156
i: [1,98,0,2,1,1,0,0] t47_subset_1
a: [3,161,0,2,1,98]
s: 161
i: [1,100,0,4,1,1,0,0] t47_subset_1
a: [3,151,0,4,1,100]
s: 151
i: [1,102,0,7,2,1,0,0] t47_subset_1
s: 102
i: [1,103,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,166,0,2,1,103]
s: 166
i: [1,104,0,2,1,0,0,255] fc1_xboole_0
a: [3,171,0,2,1,104]
s: 171
i: [1,138,0,2,1,0,0,1] rc2_zfmisc_1
a: [3,181,0,2,1,138]
s: 181
i: [1,139,0,2,1,0,0,1] rc2_xboole_0
a: [3,186,0,2,1,139]
s: 186
i: [1,140,0,2,1,0,0,1] rc1_xboole_0
a: [3,191,0,2,1,140]
s: 191
i: [1,142,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,196,0,2,1,142]
s: 196
i: [1,141,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,201,0,2,1,141]
s: 201
i: [1,106,0,3,1,0,0,1] fc2_xboole_0
s: 106
i: [1,107,0,3,1,0,0,1] fc1_subset_1
s: 107
i: [1,108,0,3,1,0,0,1] fc2_zfmisc_1
s: 108
i: [1,131,0,3,1,0,0,1] rc2_subset_1
s: 131
i: [1,105,0,3,1,0,0,255] d2_xboole_0
a: [3,176,0,3,1,105]
s: 176
r: 166
i: [1,109,0,4,2,0,0,1] cc2_zfmisc_1
s: 109
d: [2,204,1,2,1,1,40,109,171]
a: [3,210,1,2,1,204]
s: 210
i: [1,122,0,4,2,0,0,1] cc1_zfmisc_1
s: 122
i: [1,129,0,4,1,0,0,1] existence_m1_subset_1
s: 129
d: [2,206,1,2,1,1,40,109,191]
a: [3,215,1,2,1,206]
s: 215
i: [1,133,0,4,1,0,0,255] rc3_subset_1
s: 133
d: [2,205,1,3,1,0,40,109,131]
s: 205
i: [1,111,0,5,2,0,0,1] rc6_subset_1
s: 111
i: [1,115,0,5,2,0,0,1] rc1_subset_1
s: 115
i: [1,117,0,5,2,0,0,1] rc5_subset_1
s: 117
i: [1,118,0,5,2,0,0,1] rc5_subset_1
s: 118
i: [1,123,0,5,2,0,0,255] t6_boole
s: 123
i: [1,130,0,5,1,0,0,1] rc2_subset_1
s: 130
i: [1,132,0,5,1,0,0,255] rc3_subset_1
s: 132
i: [1,120,0,6,2,0,0,255] rc4_subset_1
s: 120
i: [1,110,0,7,2,0,0,1] rc6_subset_1
s: 110
i: [1,114,0,7,2,0,0,1] rc1_subset_1
s: 114
i: [1,116,0,7,2,0,0,1] rc5_subset_1
s: 116
i: [1,119,0,7,2,0,0,255] rc4_subset_1
s: 119
i: [1,128,0,7,3,0,0,1] t8_boole
s: 128
i: [1,121,0,8,3,0,0,1] cc5_subset_1
s: 121
d: [2,227,1,5,2,0,40,121,132]
s: 227
i: [1,124,0,8,3,0,0,1] cc1_subset_1
s: 124
i: [1,126,0,8,3,0,0,1] t46_subset_1
s: 126
d: [2,235,1,5,2,0,40,124,132]
s: 235
i: [1,125,0,9,3,0,0,255] cc4_subset_1
s: 125
d: [2,240,2,5,2,0,40,235,109]
s: 240
i: [1,127,0,9,3,0,0,1] t46_subset_1
s: 127
i: [1,112,0,11,4,0,0,255] cc3_subset_1
s: 112
d: [2,253,1,11,4,0,44,102,127]
s: 253
i: [1,113,0,11,4,0,0,255] cc2_subset_1
s: 113
d: [2,270,2,8,3,0,47,253]
d: [2,271,2,6,2,1,21,270,146]
d: [2,272,2,4,1,2,21,271,156]
a: [3,276,2,4,1,272]
s: 276
i: [1,137,0,14,4,0,0,255] existence_m2_subset_1
s: 137
d: [2,278,0,12,3,0,21,137,124]
s: 278
d: [2,262,1,9,3,0,40,112,132]
d: [2,269,1,5,2,0,21,262,133]
s: 269
i: [1,134,0,15,5,0,0,255] dt_m2_subset_1
s: 134
d: [2,225,1,6,2,0,40,121,129]
s: 225
d: [2,280,0,13,4,0,21,134,124]
s: 280
i: [1,135,0,15,5,0,0,255] redefinition_m2_subset_1
s: 135
d: [2,233,1,6,2,0,40,124,129]
s: 233
d: [2,283,0,13,4,0,21,135,124]
s: 283
d: [2,248,1,6,2,0,40,125,130]
s: 248
i: [1,136,0,15,5,0,0,255] redefinition_m2_subset_1
s: 136
d: [2,287,0,13,4,0,21,136,124]
s: 287
d: [2,288,1,9,3,1,40,287,151]
d: [2,297,1,7,2,2,21,288,156]
s: 297
d: [2,261,1,9,3,0,40,112,130]
d: [2,268,1,6,2,0,21,261,131]
s: 268
d: [2,302,2,12,4,2,40,297,280]
d: [2,303,2,8,3,2,21,302,151]
d: [2,304,2,6,2,2,21,303,156]
s: 304
d: [2,306,3,8,3,2,40,304,126]
d: [2,308,3,4,2,3,21,306,276]
d: [2,309,3,2,1,4,21,308,146]
d: [2,310,3,0,0,4,21,309,156]
e: 310
f: 310,272,99,100,101
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__t47_subset_1
% SZS output start Proof for subset_1__t47_subset_1
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1)) => ~(! [X2] : (m1_subset_1(X2,X0) => k1_tarski(X2) != X1) & v1_zfmisc_1(X1)))) [input t47_subset_1]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1)) => ~(! [X2] : (m1_subset_1(X2,X0) => k1_tarski(X2) != X1) & v1_zfmisc_1(X1)))) [negated conjecture 1]
4. ! [X0] : (v1_xboole_0(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_xboole_0(X1))) [input cc1_subset_1]
16. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ! [X2] : (m2_subset_1(X2,X0,X1) => m1_subset_1(X2,X0))) [input dt_m2_subset_1]
34. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ! [X2] : (m2_subset_1(X2,X0,X1) <=> m1_subset_1(X2,X1))) [input redefinition_m2_subset_1]
35. ! [X0] : ((v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) => ? [X1] : (k1_tarski(X1) = X0 & m1_subset_1(X1,X0))) [input t46_subset_1]
43. ? [X0] : (? [X1] : ((! [X2] : (k1_tarski(X2) != X1 | ~m1_subset_1(X2,X0)) & v1_zfmisc_1(X1)) & (m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1))) & ~v1_xboole_0(X0)) [ennf transformation 2]
44. ? [X0] : (? [X1] : (! [X2] : (k1_tarski(X2) != X1 | ~m1_subset_1(X2,X0)) & v1_zfmisc_1(X1) & m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [flattening 43]
57. ! [X0] : (! [X1] : (v1_xboole_0(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_xboole_0(X0)) [ennf transformation 4]
59. ! [X0] : (? [X1] : (k1_tarski(X1) = X0 & m1_subset_1(X1,X0)) | (~v1_zfmisc_1(X0) | v1_xboole_0(X0))) [ennf transformation 35]
60. ! [X0] : (? [X1] : (k1_tarski(X1) = X0 & m1_subset_1(X1,X0)) | ~v1_zfmisc_1(X0) | v1_xboole_0(X0)) [flattening 59]
62. ! [X0,X1] : (! [X2] : (m1_subset_1(X2,X0) | ~m2_subset_1(X2,X0,X1)) | (~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 16]
63. ! [X0,X1] : (! [X2] : (m1_subset_1(X2,X0) | ~m2_subset_1(X2,X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 62]
64. ! [X0,X1] : (! [X2] : (m2_subset_1(X2,X0,X1) <=> m1_subset_1(X2,X1)) | (~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 34]
65. ! [X0,X1] : (! [X2] : (m2_subset_1(X2,X0,X1) <=> m1_subset_1(X2,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 64]
68. ? [X0] : (? [X1] : (! [X2] : (k1_tarski(X2) != X1 | ~m1_subset_1(X2,X0)) & v1_zfmisc_1(X1) & m1_subset_1(X1,k1_zfmisc_1(X0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) => (? [X1] : (! [X2] : (k1_tarski(X2) != X1 | ~m1_subset_1(X2,sK0)) & v1_zfmisc_1(X1) & m1_subset_1(X1,k1_zfmisc_1(sK0)) & ~v1_xboole_0(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
69. ? [X1] : (! [X2] : (k1_tarski(X2) != X1 | ~m1_subset_1(X2,sK0)) & v1_zfmisc_1(X1) & m1_subset_1(X1,k1_zfmisc_1(sK0)) & ~v1_xboole_0(X1)) => (! [X2] : (k1_tarski(X2) != sK1 | ~m1_subset_1(X2,sK0)) & v1_zfmisc_1(sK1) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) & ~v1_xboole_0(sK1)) [choice axiom]
70. (! [X2] : (k1_tarski(X2) != sK1 | ~m1_subset_1(X2,sK0)) & v1_zfmisc_1(sK1) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) & ~v1_xboole_0(sK1)) & ~v1_xboole_0(sK0) [skolemisation 44,69,68]
79. ! [X0] : (? [X1] : (k1_tarski(X1) = X0 & m1_subset_1(X1,X0)) => (k1_tarski(sK6(X0)) = X0 & m1_subset_1(sK6(X0),X0))) [choice axiom]
80. ! [X0] : ((k1_tarski(sK6(X0)) = X0 & m1_subset_1(sK6(X0),X0)) | ~v1_zfmisc_1(X0) | v1_xboole_0(X0)) [skolemisation 60,79]
87. ! [X0,X1] : (! [X2] : ((m2_subset_1(X2,X0,X1) | ~m1_subset_1(X2,X1)) & (m1_subset_1(X2,X1) | ~m2_subset_1(X2,X0,X1))) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [nnf transformation 65]
99. ~v1_xboole_0(sK1) [cnf transformation 70] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 70] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
101. v1_zfmisc_1(sK1) [cnf transformation 70] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
102. k1_tarski(X2) != sK1 | ~m1_subset_1(X2,sK0) [cnf transformation 70] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
124. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 57] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
126. m1_subset_1(sK6(X0),X0) | ~v1_zfmisc_1(X0) | v1_xboole_0(X0) [cnf transformation 80] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
127. k1_tarski(sK6(X0)) = X0 | ~v1_zfmisc_1(X0) | v1_xboole_0(X0) [cnf transformation 80] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
134. m1_subset_1(X2,X0) | ~m2_subset_1(X2,X0,X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 63] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:1,thDist:-1}
136. m2_subset_1(X2,X0,X1) | ~m1_subset_1(X2,X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 87] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:1,thDist:-1}
144. 1 <=> v1_zfmisc_1(sK1) [avatar definition]
146. v1_zfmisc_1(sK1) <- (1) [avatar component clause 144] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
147. 1 [avatar split clause 101,144]
149. 2 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
151. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (2) [avatar component clause 149] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
152. 2 [avatar split clause 100,149]
154. 3 <=> v1_xboole_0(sK1) [avatar definition]
156. ~v1_xboole_0(sK1) <- (~3) [avatar component clause 154] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
157. ~3 [avatar split clause 99,154]
253. sK1 != X2 | ~m1_subset_1(sK6(X2),sK0) | ~v1_zfmisc_1(X2) | v1_xboole_0(X2) [superposition 102,127] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
270. ~m1_subset_1(sK6(sK1),sK0) | ~v1_zfmisc_1(sK1) | v1_xboole_0(sK1) [equality resolution 253] {a:2,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
271. ~m1_subset_1(sK6(sK1),sK0) | v1_xboole_0(sK1) <- (1) [subsumption resolution 270,146] {a:2,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
272. ~m1_subset_1(sK6(sK1),sK0) <- (1, ~3) [subsumption resolution 271,156] {a:2,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
274. 15 <=> m1_subset_1(sK6(sK1),sK0) [avatar definition]
276. ~m1_subset_1(sK6(sK1),sK0) <- (~15) [avatar component clause 274] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
277. ~15 | ~1 | 3 [avatar split clause 272,154,144,274]
280. ~m2_subset_1(X2,X0,X1) | m1_subset_1(X2,X0) | ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_xboole_0(X1) [subsumption resolution 134,124] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
287. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | ~m1_subset_1(X2,X1) | m2_subset_1(X2,X0,X1) | v1_xboole_0(X1) [subsumption resolution 136,124] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
288. ~m1_subset_1(X0,sK1) | m2_subset_1(X0,sK0,sK1) | v1_xboole_0(sK1) <- (2) [resolution 287,151] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
297. m2_subset_1(X0,sK0,sK1) | ~m1_subset_1(X0,sK1) <- (2, ~3) [subsumption resolution 288,156] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
302. ~m1_subset_1(X1,sK1) | m1_subset_1(X1,sK0) | ~m1_subset_1(sK1,k1_zfmisc_1(sK0)) | v1_xboole_0(sK1) <- (2, ~3) [resolution 297,280] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
303. ~m1_subset_1(X1,sK1) | m1_subset_1(X1,sK0) | v1_xboole_0(sK1) <- (2, ~3) [subsumption resolution 302,151] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
304. ~m1_subset_1(X1,sK1) | m1_subset_1(X1,sK0) <- (2, ~3) [subsumption resolution 303,156] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
306. m1_subset_1(sK6(sK1),sK0) | ~v1_zfmisc_1(sK1) | v1_xboole_0(sK1) <- (2, ~3) [resolution 304,126] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
308. ~v1_zfmisc_1(sK1) | v1_xboole_0(sK1) <- (2, ~3, ~15) [subsumption resolution 306,276] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
309. v1_xboole_0(sK1) <- (1, 2, ~3, ~15) [subsumption resolution 308,146] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
310. $false <- (1, 2, ~3, ~15) [subsumption resolution 309,156] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
311. ~1 | ~2 | 3 | 15 [avatar contradiction clause 310]
312. $false [avatar sat refutation 147,152,157,277,311] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:61,thDist:-61}
% SZS output end Proof for subset_1__t47_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 37
% >>> Preprocessing
% Introduced skolems: 15
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 45
% Generated clauses: 153
% Active clauses: 58
% Passive clauses: 84
% Final active clauses: 57
% Final passive clauses: 20
% Main loop iterations started: 64
% 
% >>> Simplifying Inferences
% Duplicate literals: 7
% Fw subsumption resolutions: 22
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 18
% 
% >>> Generating Inferences
% Binary resolution: 58
% Backward superposition: 5
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 12
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.001 s ( own 0.0000 s ) 

