i: [1,234,0,2,1,1,0,0] t17_tops_3
a: [3,347,0,2,1,234]
s: 347
i: [1,233,0,2,1,1,0,0] t17_tops_3
a: [3,352,0,2,1,233]
s: 352
i: [1,232,0,2,1,1,0,0] t17_tops_3
a: [3,357,0,2,1,232]
s: 357
i: [1,237,0,3,1,1,0,0] t17_tops_3
a: [3,332,0,3,1,237]
s: 332
i: [1,236,0,3,1,1,0,0] t17_tops_3
a: [3,337,0,3,1,236]
s: 337
d: [2,384,0,3,1,2,22,337,332]
a: [3,388,0,3,1,384]
s: 388
i: [1,235,0,5,1,1,0,0] t17_tops_3
a: [3,342,0,5,1,235]
s: 342
d: [2,390,0,5,1,2,22,342,332]
a: [3,394,0,5,1,390]
s: 394
i: [1,238,0,2,1,0,0,255] spc1_boole
a: [3,362,0,2,1,238]
s: 362
i: [1,239,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,367,0,2,1,239]
s: 367
i: [1,327,0,2,1,0,0,2] existence_l1_struct_0
a: [3,377,0,2,1,327]
s: 377
i: [1,328,0,2,1,0,0,1] existence_l1_pre_topc
a: [3,382,0,2,1,328]
s: 382
i: [1,322,0,3,1,0,0,255] reflexivity_r1_tarski
s: 322
i: [1,240,0,3,1,0,0,255] d2_xboole_0
a: [3,372,0,3,1,240]
s: 372
r: 367
d: [2,396,0,2,1,2,23,367,372]
a: [3,400,0,2,1,396]
s: 400
i: [1,257,0,4,2,0,0,2] dt_l1_pre_topc
s: 257
d: [2,402,1,2,1,1,40,257,347]
a: [3,407,1,2,1,402]
s: 407
i: [1,321,0,4,1,0,0,1] existence_m1_subset_1
s: 321
d: [2,403,1,2,1,1,40,257,382]
a: [3,412,1,2,1,403]
s: 412
i: [1,268,0,5,2,0,0,1] t6_boole
s: 268
i: [1,318,0,5,2,0,0,255] t7_boole
s: 318
d: [2,417,1,5,2,1,44,388,268]
s: 417
i: [1,243,0,6,3,0,0,255] cc7_struct_0
s: 243
d: [2,416,1,7,2,1,44,394,268]
s: 416
i: [1,244,0,6,3,0,0,255] cc5_struct_0
s: 244
i: [1,246,0,6,3,0,0,255] cc2_struct_0
s: 246
d: [2,418,1,4,2,1,40,244,377]
a: [3,424,1,2,1,418]
s: 424
i: [1,247,0,6,3,0,0,255] cc17_struct_0
s: 247
d: [2,420,2,4,2,1,40,244,412]
a: [3,433,2,2,1,420]
s: 433
i: [1,248,0,6,3,0,0,255] cc6_struct_0
s: 248
i: [1,249,0,6,3,0,0,255] cc1_struct_0
s: 249
i: [1,252,0,6,3,0,0,255] cc4_struct_0
s: 252
i: [1,258,0,6,2,0,0,1] fc12_tops_1
s: 258
i: [1,265,0,6,2,0,0,1] rc4_tops_1
s: 265
i: [1,267,0,6,2,0,0,255] rc5_tops_1
s: 267
i: [1,314,0,6,2,0,0,255] t1_subset
s: 314
i: [1,323,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 323
i: [1,241,0,7,2,0,0,1] d3_struct_0
s: 241
d: [2,452,1,8,3,0,44,258,241]
d: [2,453,1,6,2,0,21,452,257]
s: 453
i: [1,250,0,7,3,0,0,255] cc8_struct_0
s: 250
i: [1,253,0,7,3,0,0,255] cc11_struct_0
s: 253
i: [1,254,0,7,3,0,0,255] cc11_struct_0
s: 254
i: [1,255,0,7,3,0,0,255] cc9_struct_0
s: 255
i: [1,269,0,7,3,0,0,255] fc9_struct_0
s: 269
i: [1,270,0,7,3,0,0,255] fc6_struct_0
s: 270
i: [1,272,0,7,3,0,0,255] rc22_struct_0
s: 272
i: [1,284,0,7,3,0,0,2] fc5_struct_0
s: 284
i: [1,285,0,7,3,0,0,1] fc2_struct_0
s: 285
i: [1,287,0,7,3,0,0,1] rc4_struct_0
s: 287
i: [1,289,0,7,3,0,0,255] rc21_struct_0
s: 289
i: [1,290,0,7,3,0,0,255] rc21_struct_0
s: 290
i: [1,310,0,7,3,0,0,255] fc8_struct_0
s: 310
i: [1,311,0,7,3,0,0,255] fc7_struct_0
s: 311
i: [1,312,0,7,3,0,0,2] fc4_struct_0
s: 312
i: [1,313,0,7,3,0,0,1] fc1_struct_0
s: 313
i: [1,316,0,7,2,0,0,255] t3_subset
s: 316
i: [1,317,0,7,3,0,0,1] t8_boole
s: 317
i: [1,242,0,8,2,0,0,2] dt_k2_struct_0
s: 242
d: [2,414,1,7,3,0,44,268,268]
s: 414
i: [1,264,0,8,2,0,0,1] rc4_tops_1
s: 264
i: [1,266,0,8,2,0,0,255] rc5_tops_1
s: 266
d: [2,458,1,9,4,0,44,284,241]
d: [2,459,1,7,3,0,18,458]
s: 459
i: [1,291,0,8,3,0,0,255] fc10_tops_1
s: 291
d: [2,465,1,9,4,0,44,312,241]
d: [2,466,1,7,3,0,18,465]
s: 466
i: [1,292,0,8,3,0,0,3] fc4_pre_topc
s: 292
i: [1,300,0,8,3,0,0,255] rc7_tops_1
s: 300
i: [1,302,0,8,3,0,0,255] rc1_tops_1
s: 302
i: [1,304,0,8,3,0,0,255] rc10_pre_topc
s: 304
i: [1,306,0,8,3,0,0,255] rc2_tops_1
s: 306
i: [1,307,0,8,3,0,0,255] rc2_tops_1
s: 307
i: [1,309,0,8,3,0,0,3] rc6_pre_topc
s: 309
i: [1,315,0,8,3,0,0,255] t2_subset
s: 315
i: [1,256,0,9,4,0,0,255] cc10_struct_0
s: 256
d: [2,476,2,11,3,1,40,315,416]
a: [3,489,2,4,1,476]
s: 489
i: [1,275,0,9,4,0,0,255] rc6_tops_1
s: 275
i: [1,278,0,9,4,0,0,255] rc3_tops_1
s: 278
d: [2,482,1,6,2,0,40,315,321]
s: 482
i: [1,282,0,9,4,0,0,3] rc7_pre_topc
s: 282
d: [2,500,2,6,2,0,40,482,323]
s: 500
i: [1,320,0,9,3,0,0,255] t5_subset
s: 320
i: [1,271,0,10,3,0,0,255] rc22_struct_0
s: 271
d: [2,505,1,6,2,1,40,320,394]
a: [3,516,1,3,1,505]
s: 516
i: [1,273,0,10,4,0,0,255] fc14_tops_1
s: 273
d: [2,509,1,7,2,0,40,320,321]
s: 509
i: [1,276,0,10,4,0,0,255] rc6_tops_1
s: 276
i: [1,279,0,10,4,0,0,255] rc3_tops_1
s: 279
d: [2,523,2,6,2,0,40,509,482]
s: 523
i: [1,280,0,10,4,0,0,255] rc3_tops_1
s: 280
d: [2,470,1,10,3,0,44,242,241]
d: [2,471,1,8,2,0,18,470]
s: 471
i: [1,283,0,10,4,0,0,3] rc7_pre_topc
s: 283
i: [1,286,0,10,3,0,0,1] rc4_struct_0
s: 286
d: [2,472,1,10,4,0,44,291,241]
d: [2,473,1,8,3,0,21,472,257]
s: 473
i: [1,288,0,10,3,0,0,255] rc21_struct_0
s: 288
d: [2,474,1,10,4,0,44,292,241]
d: [2,475,1,8,3,0,21,474,257]
s: 475
i: [1,299,0,10,3,0,0,255] rc7_tops_1
s: 299
i: [1,301,0,10,3,0,0,255] rc1_tops_1
s: 301
d: [2,504,1,8,3,0,40,320,316]
s: 504
i: [1,303,0,10,3,0,0,255] rc10_pre_topc
s: 303
d: [2,454,1,9,4,0,44,250,268]
s: 454
i: [1,305,0,10,3,0,0,255] rc2_tops_1
s: 305
i: [1,308,0,10,3,0,0,3] rc6_pre_topc
s: 308
d: [2,456,1,9,4,0,44,255,268]
s: 456
i: [1,319,0,10,3,0,0,255] t4_subset
s: 319
d: [2,556,1,7,2,1,40,319,394]
s: 556
i: [1,261,0,12,4,0,0,255] cc2_tops_1
s: 261
i: [1,274,0,12,4,0,0,255] rc6_tops_1
s: 274
d: [2,574,1,7,3,1,40,261,394]
d: [2,597,1,5,2,2,21,574,400]
d: [2,598,1,3,1,3,21,597,347]
a: [3,602,1,3,1,598]
s: 602
i: [1,277,0,12,4,0,0,255] rc3_tops_1
s: 277
d: [2,571,2,9,4,1,40,261,416]
d: [2,594,2,7,3,1,18,571]
d: [2,595,2,5,2,2,21,594,347]
s: 595
i: [1,281,0,12,4,0,0,3] rc7_pre_topc
s: 281
i: [1,293,0,14,5,0,0,255] cc3_tops_1
s: 293
d: [2,635,1,9,4,1,40,293,394]
d: [2,665,1,7,3,2,21,635,400]
d: [2,666,1,5,2,3,21,665,347]
d: [2,667,1,3,1,4,21,666,352]
a: [3,671,1,3,1,667]
s: 671
i: [1,294,0,14,5,0,0,255] cc1_tops_1
s: 294
d: [2,679,1,9,4,1,40,294,394]
d: [2,708,1,7,3,2,21,679,400]
d: [2,709,1,5,2,3,21,708,347]
d: [2,710,1,3,1,4,21,709,352]
a: [3,714,1,3,1,710]
s: 714
i: [1,295,0,14,5,0,0,3] cc2_pre_topc
s: 295
i: [1,324,0,14,4,0,0,2] fc1_tops_1
s: 324
d: [2,721,1,9,4,1,40,295,394]
d: [2,750,1,7,3,2,21,721,400]
d: [2,751,1,5,2,3,21,750,347]
d: [2,752,1,3,1,4,21,751,352]
a: [3,756,1,3,1,752]
s: 756
i: [1,325,0,14,3,0,0,2] dt_k2_pre_topc
s: 325
d: [2,617,3,9,4,2,40,595,273]
d: [2,619,3,7,3,2,21,617,347]
d: [2,620,3,5,2,3,21,619,352]
d: [2,621,3,3,1,4,21,620,357]
a: [3,625,3,3,1,621]
s: 625
i: [1,259,0,15,4,0,0,2] t22_pre_topc
s: 259
i: [1,296,0,15,5,0,0,255] cc4_tops_1
s: 296
d: [2,762,1,9,3,1,40,324,394]
d: [2,791,1,7,2,2,21,762,347]
d: [2,792,1,5,1,3,21,791,352]
a: [3,796,1,5,1,792]
s: 796
i: [1,262,0,16,4,0,0,1] d3_tops_1
s: 262
d: [2,906,2,14,4,1,44,796,262]
d: [2,908,2,11,3,2,21,906,388]
d: [2,909,2,6,2,3,21,908,394]
d: [2,910,2,4,1,4,21,909,347]
a: [3,914,2,4,1,910]
s: 914
i: [1,263,0,16,4,0,0,1] d3_tops_1
s: 263
i: [1,326,0,16,3,0,0,2] projectivity_k2_pre_topc
s: 326
d: [2,916,3,6,2,1,44,914,241]
d: [2,917,3,4,1,2,21,916,407]
a: [3,921,3,4,1,917]
s: 921
i: [1,260,0,17,5,0,0,2] t22_pre_topc
s: 260
d: [2,821,1,10,3,1,40,259,394]
d: [2,851,1,7,2,2,21,821,756]
d: [2,852,1,5,1,3,21,851,347]
a: [3,856,1,5,1,852]
s: 856
r: 796
d: [2,951,2,14,4,1,44,262,856]
d: [2,962,2,11,3,2,21,951,388]
d: [2,963,2,6,2,3,21,962,394]
d: [2,964,2,4,1,4,21,963,347]
f: 964,852,752,621,396,390,384,240,239,232,233,234,235,236,237
% Refutation found. Thanks to Tanya!
% SZS status Theorem for tops_3__t17_tops_3
% SZS output start Proof for tops_3__t17_tops_3
1. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ~(k1_xboole_0 = X1 & v1_tops_1(X1,X0)))) [input t17_tops_3]
2. ~! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => ~(k1_xboole_0 = X1 & v1_tops_1(X1,X0)))) [negated conjecture 1]
9. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v1_xboole_0(X1) => v4_pre_topc(X1,X0)))) [input cc2_pre_topc]
11. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v1_xboole_0(X1) => v2_tops_1(X1,X0)))) [input cc2_tops_1]
23. k1_xboole_0 = o_0_0_xboole_0 [input d2_xboole_0]
25. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (v1_tops_1(X1,X0) <=> k2_struct_0(X0) = k2_pre_topc(X0,X1)))) [input d3_tops_1]
33. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
41. ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => ~v2_tops_1(k2_struct_0(X0),X0)) [input fc14_tops_1]
69. ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) => (((k2_pre_topc(X0,X1) = X1 & v2_pre_topc(X0)) => v4_pre_topc(X1,X0)) & (v4_pre_topc(X1,X0) => k2_pre_topc(X0,X1) = X1)))) [input t22_pre_topc]
74. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
88. ? [X0] : (? [X1] : ((k1_xboole_0 = X1 & v1_tops_1(X1,X0)) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & (l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
89. ? [X0] : (? [X1] : (k1_xboole_0 = X1 & v1_tops_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) [flattening 88]
116. ! [X0] : (! [X1] : (((v4_pre_topc(X1,X0) | (k2_pre_topc(X0,X1) != X1 | ~v2_pre_topc(X0))) & (k2_pre_topc(X0,X1) = X1 | ~v4_pre_topc(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 69]
117. ! [X0] : (! [X1] : (((v4_pre_topc(X1,X0) | k2_pre_topc(X0,X1) != X1 | ~v2_pre_topc(X0)) & (k2_pre_topc(X0,X1) = X1 | ~v4_pre_topc(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [flattening 116]
118. ! [X0] : (! [X1] : ((v2_tops_1(X1,X0) | ~v1_xboole_0(X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 11]
119. ! [X0] : (! [X1] : (v2_tops_1(X1,X0) | ~v1_xboole_0(X1) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [flattening 118]
120. ! [X0] : (! [X1] : ((v1_tops_1(X1,X0) <=> k2_struct_0(X0) = k2_pre_topc(X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [ennf transformation 25]
123. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 74]
130. ! [X0] : (~v2_tops_1(k2_struct_0(X0),X0) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0))) [ennf transformation 41]
131. ! [X0] : (~v2_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0)) [flattening 130]
154. ! [X0] : (! [X1] : ((v4_pre_topc(X1,X0) | ~v1_xboole_0(X1)) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | (~l1_pre_topc(X0) | ~v2_pre_topc(X0))) [ennf transformation 9]
155. ! [X0] : (! [X1] : (v4_pre_topc(X1,X0) | ~v1_xboole_0(X1) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0)) [flattening 154]
196. ? [X0] : (? [X1] : (k1_xboole_0 = X1 & v1_tops_1(X1,X0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~v2_struct_0(X0)) => (? [X1] : (k1_xboole_0 = X1 & v1_tops_1(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0)) [choice axiom]
197. ? [X1] : (k1_xboole_0 = X1 & v1_tops_1(X1,sK0) & m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(sK0)))) => (k1_xboole_0 = sK1 & v1_tops_1(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) [choice axiom]
198. (k1_xboole_0 = sK1 & v1_tops_1(sK1,sK0) & m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0)))) & l1_pre_topc(sK0) & v2_pre_topc(sK0) & ~v2_struct_0(sK0) [skolemisation 89,197,196]
199. ! [X0] : (! [X1] : (((v1_tops_1(X1,X0) | k2_struct_0(X0) != k2_pre_topc(X0,X1)) & (k2_struct_0(X0) = k2_pre_topc(X0,X1) | ~v1_tops_1(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0)))) | ~l1_pre_topc(X0)) [nnf transformation 120]
232. ~v2_struct_0(sK0) [cnf transformation 198] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
233. v2_pre_topc(sK0) [cnf transformation 198] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
234. l1_pre_topc(sK0) [cnf transformation 198] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
235. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [cnf transformation 198] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
236. v1_tops_1(sK1,sK0) [cnf transformation 198] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
237. k1_xboole_0 = sK1 [cnf transformation 198] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
239. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 33] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
240. k1_xboole_0 = o_0_0_xboole_0 [cnf transformation 23] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
259. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v4_pre_topc(X1,X0) | k2_pre_topc(X0,X1) = X1 | ~l1_pre_topc(X0) [cnf transformation 117] {a:0,w:15,wCS:6000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
261. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v1_xboole_0(X1) | v2_tops_1(X1,X0) | ~l1_pre_topc(X0) [cnf transformation 119] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
262. k2_struct_0(X0) = k2_pre_topc(X0,X1) | ~v1_tops_1(X1,X0) | ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~l1_pre_topc(X0) [cnf transformation 199] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
268. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 123] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
273. ~v2_tops_1(k2_struct_0(X0),X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) | v2_struct_0(X0) [cnf transformation 131] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
295. ~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X0))) | ~v1_xboole_0(X1) | v4_pre_topc(X1,X0) | ~l1_pre_topc(X0) | ~v2_pre_topc(X0) [cnf transformation 155] {a:0,w:14,wCS:5600,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
330. 1 <=> k1_xboole_0 = sK1 [avatar definition]
332. k1_xboole_0 = sK1 <- (1) [avatar component clause 330] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
333. 1 [avatar split clause 237,330]
335. 2 <=> v1_tops_1(sK1,sK0) [avatar definition]
337. v1_tops_1(sK1,sK0) <- (2) [avatar component clause 335] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
338. 2 [avatar split clause 236,335]
340. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
342. m1_subset_1(sK1,k1_zfmisc_1(u1_struct_0(sK0))) <- (3) [avatar component clause 340] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
343. 3 [avatar split clause 235,340]
345. 4 <=> l1_pre_topc(sK0) [avatar definition]
347. l1_pre_topc(sK0) <- (4) [avatar component clause 345] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
348. 4 [avatar split clause 234,345]
350. 5 <=> v2_pre_topc(sK0) [avatar definition]
352. v2_pre_topc(sK0) <- (5) [avatar component clause 350] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
353. 5 [avatar split clause 233,350]
355. 6 <=> v2_struct_0(sK0) [avatar definition]
357. ~v2_struct_0(sK0) <- (~6) [avatar component clause 355] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
358. ~6 [avatar split clause 232,355]
365. 8 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
367. v1_xboole_0(o_0_0_xboole_0) <- (8) [avatar component clause 365] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
368. 8 [avatar split clause 239,365]
370. 9 <=> k1_xboole_0 = o_0_0_xboole_0 [avatar definition]
372. k1_xboole_0 = o_0_0_xboole_0 <- (9) [avatar component clause 370] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
373. 9 [avatar split clause 240,370]
384. v1_tops_1(k1_xboole_0,sK0) <- (1, 2) [forward demodulation 337,332] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
386. 12 <=> v1_tops_1(k1_xboole_0,sK0) [avatar definition]
388. v1_tops_1(k1_xboole_0,sK0) <- (12) [avatar component clause 386] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
389. 12 | ~1 | ~2 [avatar split clause 384,335,330,386]
390. m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) <- (1, 3) [forward demodulation 342,332] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
392. 13 <=> m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) [avatar definition]
394. m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) <- (13) [avatar component clause 392] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
395. 13 | ~1 | ~3 [avatar split clause 390,340,330,392]
396. v1_xboole_0(k1_xboole_0) <- (8, 9) [backward demodulation 367,372] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
398. 14 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
400. v1_xboole_0(k1_xboole_0) <- (14) [avatar component clause 398] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
401. 14 | ~8 | ~9 [avatar split clause 396,370,365,398]
416. m1_subset_1(X0,k1_zfmisc_1(u1_struct_0(sK0))) | ~v1_xboole_0(X0) <- (13) [superposition 394,268] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
571. ~v1_xboole_0(X0) | v2_tops_1(X0,sK0) | ~l1_pre_topc(sK0) | ~v1_xboole_0(X0) <- (13) [resolution 261,416] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
594. ~v1_xboole_0(X0) | v2_tops_1(X0,sK0) | ~l1_pre_topc(sK0) <- (13) [duplicate literal removal 571] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
595. v2_tops_1(X0,sK0) | ~v1_xboole_0(X0) <- (4, 13) [subsumption resolution 594,347] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
617. ~v1_xboole_0(k2_struct_0(sK0)) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 13) [resolution 595,273] {a:3,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
619. ~v1_xboole_0(k2_struct_0(sK0)) | ~v2_pre_topc(sK0) | v2_struct_0(sK0) <- (4, 13) [subsumption resolution 617,347] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
620. ~v1_xboole_0(k2_struct_0(sK0)) | v2_struct_0(sK0) <- (4, 5, 13) [subsumption resolution 619,352] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
621. ~v1_xboole_0(k2_struct_0(sK0)) <- (4, 5, ~6, 13) [subsumption resolution 620,357] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
623. 29 <=> v1_xboole_0(k2_struct_0(sK0)) [avatar definition]
626. ~29 | ~4 | ~5 | 6 | ~13 [avatar split clause 621,392,355,350,345,623]
721. ~v1_xboole_0(k1_xboole_0) | v4_pre_topc(k1_xboole_0,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (13) [resolution 295,394] {a:1,w:9,wCS:900,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
750. v4_pre_topc(k1_xboole_0,sK0) | ~l1_pre_topc(sK0) | ~v2_pre_topc(sK0) <- (13, 14) [subsumption resolution 721,400] {a:1,w:7,wCS:700,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
751. v4_pre_topc(k1_xboole_0,sK0) | ~v2_pre_topc(sK0) <- (4, 13, 14) [subsumption resolution 750,347] {a:1,w:5,wCS:500,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
752. v4_pre_topc(k1_xboole_0,sK0) <- (4, 5, 13, 14) [subsumption resolution 751,352] {a:1,w:3,wCS:300,goal:1,sine:3,thAx:0,allAx:4,thDist:-4}
754. 32 <=> v4_pre_topc(k1_xboole_0,sK0) [avatar definition]
756. v4_pre_topc(k1_xboole_0,sK0) <- (32) [avatar component clause 754] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
757. 32 | ~4 | ~5 | ~13 | ~14 [avatar split clause 752,398,392,350,345,754]
821. ~v4_pre_topc(k1_xboole_0,sK0) | k1_xboole_0 = k2_pre_topc(sK0,k1_xboole_0) | ~l1_pre_topc(sK0) <- (13) [resolution 259,394] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
851. k1_xboole_0 = k2_pre_topc(sK0,k1_xboole_0) | ~l1_pre_topc(sK0) <- (13, 32) [subsumption resolution 821,756] {a:1,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
852. k1_xboole_0 = k2_pre_topc(sK0,k1_xboole_0) <- (4, 13, 32) [subsumption resolution 851,347] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
854. 34 <=> k1_xboole_0 = k2_pre_topc(sK0,k1_xboole_0) [avatar definition]
856. k1_xboole_0 = k2_pre_topc(sK0,k1_xboole_0) <- (34) [avatar component clause 854] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
857. 34 | ~4 | ~13 | ~32 [avatar split clause 852,754,392,345,854]
951. k1_xboole_0 = k2_struct_0(sK0) | ~v1_tops_1(k1_xboole_0,sK0) | ~m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (34) [superposition 262,856] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
958. 37 <=> k1_xboole_0 = k2_struct_0(sK0) [avatar definition]
962. k1_xboole_0 = k2_struct_0(sK0) | ~m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(sK0))) | ~l1_pre_topc(sK0) <- (12, 34) [subsumption resolution 951,388] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
963. k1_xboole_0 = k2_struct_0(sK0) | ~l1_pre_topc(sK0) <- (12, 13, 34) [subsumption resolution 962,394] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
964. k1_xboole_0 = k2_struct_0(sK0) <- (4, 12, 13, 34) [subsumption resolution 963,347] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
965. 37 | ~4 | ~12 | ~13 | ~34 [avatar split clause 964,854,392,386,345,958]
966. k1_xboole_0 != o_0_0_xboole_0 | k1_xboole_0 != k2_struct_0(sK0) | v1_xboole_0(k2_struct_0(sK0)) | ~v1_xboole_0(o_0_0_xboole_0) [theory tautology sat conflict] {a:0,w:12,wCS:4800,sine:255,thAx:0,allAx:1,thDist:-1}
967. $false [avatar sat refutation 333,338,343,348,353,358,368,373,389,395,401,626,757,857,965,966] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:213,thDist:-213}
% SZS output end Proof for tops_3__t17_tops_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 75
% >>> Preprocessing
% Introduced skolems: 18
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 97
% Generated clauses: 583
% Active clauses: 120
% Passive clauses: 290
% Final active clauses: 118
% Final passive clauses: 161
% Main loop iterations started: 129
% 
% >>> Simplifying Inferences
% Duplicate literals: 154
% Trivial inequalities: 2
% Fw subsumption resolutions: 75
% Fw demodulations: 2
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Forward subsumptions: 65
% 
% >>> Generating Inferences
% Binary resolution: 231
% Forward superposition: 22
% Backward superposition: 14
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 14
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 37
% SAT solver unit clauses: 11
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9850
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s ( own 0.0000 s ) 
% sine selection: 0.001 s
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

