i: [1,58,0,2,1,1,0,0] t3_algstr_2
a: [3,89,0,2,1,58]
s: 89
i: [1,57,0,2,1,1,0,0] t3_algstr_2
a: [3,94,0,2,1,57]
s: 94
i: [1,56,0,2,1,1,0,0] t3_algstr_2
a: [3,99,0,2,1,56]
s: 99
i: [1,55,0,2,1,1,0,0] t3_algstr_2
a: [3,104,0,2,1,55]
s: 104
i: [1,54,0,2,1,1,0,0] t3_algstr_2
a: [3,109,0,2,1,54]
s: 109
i: [1,59,0,4,1,1,0,0] t3_algstr_2
a: [3,84,0,4,1,59]
s: 84
i: [1,60,0,7,1,1,0,0] t3_algstr_2
a: [3,79,0,7,1,60]
s: 79
i: [1,72,0,2,1,0,0,255] existence_l2_struct_0
a: [3,114,0,2,1,72]
s: 114
i: [1,73,0,2,1,0,0,1] existence_l2_algstr_0
a: [3,119,0,2,1,73]
s: 119
i: [1,74,0,2,1,0,0,255] existence_l1_algstr_0
a: [3,124,0,2,1,74]
s: 124
i: [1,62,0,4,2,0,0,255] dt_l2_algstr_0
s: 62
d: [2,126,1,2,1,1,40,62,89]
a: [3,131,1,2,1,126]
s: 131
i: [1,63,0,4,2,0,0,255] dt_l2_algstr_0
s: 63
d: [2,138,1,2,1,1,40,63,89]
a: [3,143,1,2,1,138]
s: 143
i: [1,66,0,4,1,0,0,1] existence_m1_subset_1
s: 66
i: [1,61,0,7,2,0,0,255] dt_k4_struct_0
s: 61
d: [2,127,1,2,1,1,40,62,119]
a: [3,136,1,2,1,127]
s: 136
i: [1,71,0,17,4,0,0,255] dt_k1_algstr_0
s: 71
d: [2,139,1,2,1,1,40,63,119]
a: [3,148,1,2,1,139]
s: 148
i: [1,67,0,18,6,0,0,1] dt_k1_algstr_2
s: 67
i: [1,68,0,19,5,0,0,255] dt_k3_rlvect_1
s: 68
i: [1,69,0,21,5,0,0,255] redefinition_k3_rlvect_1
s: 69
i: [1,70,0,21,5,0,0,255] commutativity_k3_rlvect_1
s: 70
i: [1,75,0,27,7,0,0,255] d1_algstr_2
s: 75
d: [2,156,0,21,6,0,21,75,67]
s: 156
i: [1,65,0,28,8,0,0,255] d1_algstr_2
s: 65
d: [2,159,1,29,9,0,44,71,156]
d: [2,160,1,25,8,0,18,159]
d: [2,167,1,23,7,0,21,160,63]
d: [2,168,1,17,6,0,21,167,67]
s: 168
d: [2,157,1,35,10,0,44,156,69]
d: [2,162,1,31,9,0,18,157]
d: [2,163,1,29,8,0,21,162,63]
d: [2,164,1,23,7,0,21,163,67]
s: 164
d: [2,158,1,35,10,0,44,69,156]
d: [2,161,1,31,9,0,18,158]
d: [2,165,1,29,8,0,21,161,63]
d: [2,166,1,23,7,0,21,165,67]
s: 166
d: [2,169,1,40,12,0,44,65,69]
d: [2,171,1,32,10,0,18,169]
d: [2,172,1,30,9,0,21,171,63]
s: 172
d: [2,175,2,42,13,0,44,172,70]
d: [2,179,2,32,10,0,18,175]
d: [2,180,2,30,9,0,21,179,63]
s: 180
d: [2,184,3,46,15,0,44,180,164]
d: [2,185,3,32,9,0,18,184]
d: [2,186,3,27,8,0,19,185]
d: [2,189,3,21,7,0,21,186,67]
s: 189
d: [2,190,4,0,0,7,52,94,109,104,99,89,84,79,189]
e: 190
d: [2,191,4,17,7,1,44,79,189]
d: [2,198,4,14,6,1,19,191]
d: [2,200,4,10,5,2,21,198,84]
d: [2,201,4,8,4,3,21,200,89]
d: [2,202,4,6,3,4,21,201,99]
d: [2,203,4,4,2,5,21,202,104]
d: [2,204,4,2,1,6,21,203,109]
d: [2,205,4,0,0,7,21,204,94]
e: 205
f: 205,54,55,56,57,58,59,60
% Refutation found. Thanks to Tanya!
% SZS status Theorem for algstr_2__t3_algstr_2
% SZS output start Proof for algstr_2__t3_algstr_2
1. ! [X0] : ((l2_algstr_0(X0) & v2_rlvect_1(X0) & v3_algstr_1(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k1_algstr_2(X0,k1_algstr_2(X0,X1)) = X1)) [input t3_algstr_2]
2. ~! [X0] : ((l2_algstr_0(X0) & v2_rlvect_1(X0) & v3_algstr_1(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k1_algstr_2(X0,k1_algstr_2(X0,X1)) = X1)) [negated conjecture 1]
4. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_algstr_0(X0) & v2_rlvect_1(X0)) => k3_rlvect_1(X0,X1,X2) = k3_rlvect_1(X0,X2,X1)) [input commutativity_k3_rlvect_1]
5. ! [X0] : ((l2_algstr_0(X0) & v3_algstr_1(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_algstr_2(X0,X1) = X2 <=> k1_algstr_0(X0,X1,X2) = k4_struct_0(X0))))) [input d1_algstr_2]
7. ! [X0,X1] : ((m1_subset_1(X1,u1_struct_0(X0)) & l2_algstr_0(X0) & v3_algstr_1(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) => m1_subset_1(k1_algstr_2(X0,X1),u1_struct_0(X0))) [input dt_k1_algstr_2]
12. ! [X0] : (l2_algstr_0(X0) => (l1_algstr_0(X0) & l2_struct_0(X0))) [input dt_l2_algstr_0]
21. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_algstr_0(X0) & v2_rlvect_1(X0)) => k3_rlvect_1(X0,X1,X2) = k1_algstr_0(X0,X1,X2)) [input redefinition_k3_rlvect_1]
26. ? [X0] : (? [X1] : (k1_algstr_2(X0,k1_algstr_2(X0,X1)) != X1 & m1_subset_1(X1,u1_struct_0(X0))) & (l2_algstr_0(X0) & v2_rlvect_1(X0) & v3_algstr_1(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
27. ? [X0] : (? [X1] : (k1_algstr_2(X0,k1_algstr_2(X0,X1)) != X1 & m1_subset_1(X1,u1_struct_0(X0))) & l2_algstr_0(X0) & v2_rlvect_1(X0) & v3_algstr_1(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) [flattening 26]
29. ! [X0] : ((l1_algstr_0(X0) & l2_struct_0(X0)) | ~l2_algstr_0(X0)) [ennf transformation 12]
30. ! [X0] : (! [X1] : (! [X2] : ((k1_algstr_2(X0,X1) = X2 <=> k1_algstr_0(X0,X1,X2) = k4_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 5]
31. ! [X0] : (! [X1] : (! [X2] : ((k1_algstr_2(X0,X1) = X2 <=> k1_algstr_0(X0,X1,X2) = k4_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0)) [flattening 30]
32. ! [X0,X1] : (m1_subset_1(k1_algstr_2(X0,X1),u1_struct_0(X0)) | (~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 7]
33. ! [X0,X1] : (m1_subset_1(k1_algstr_2(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0)) [flattening 32]
36. ! [X0,X1,X2] : (k3_rlvect_1(X0,X1,X2) = k1_algstr_0(X0,X1,X2) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_algstr_0(X0) | ~v2_rlvect_1(X0))) [ennf transformation 21]
37. ! [X0,X1,X2] : (k3_rlvect_1(X0,X1,X2) = k1_algstr_0(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_algstr_0(X0) | ~v2_rlvect_1(X0)) [flattening 36]
38. ! [X0,X1,X2] : (k3_rlvect_1(X0,X1,X2) = k3_rlvect_1(X0,X2,X1) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_algstr_0(X0) | ~v2_rlvect_1(X0))) [ennf transformation 4]
39. ! [X0,X1,X2] : (k3_rlvect_1(X0,X1,X2) = k3_rlvect_1(X0,X2,X1) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_algstr_0(X0) | ~v2_rlvect_1(X0)) [flattening 38]
42. ? [X0] : (? [X1] : (k1_algstr_2(X0,k1_algstr_2(X0,X1)) != X1 & m1_subset_1(X1,u1_struct_0(X0))) & l2_algstr_0(X0) & v2_rlvect_1(X0) & v3_algstr_1(X0) & v5_algstr_0(X0) & ~v2_struct_0(X0)) => (? [X1] : (k1_algstr_2(sK0,k1_algstr_2(sK0,X1)) != X1 & m1_subset_1(X1,u1_struct_0(sK0))) & l2_algstr_0(sK0) & v2_rlvect_1(sK0) & v3_algstr_1(sK0) & v5_algstr_0(sK0) & ~v2_struct_0(sK0)) [choice axiom]
43. ? [X1] : (k1_algstr_2(sK0,k1_algstr_2(sK0,X1)) != X1 & m1_subset_1(X1,u1_struct_0(sK0))) => (sK1 != k1_algstr_2(sK0,k1_algstr_2(sK0,sK1)) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
44. (sK1 != k1_algstr_2(sK0,k1_algstr_2(sK0,sK1)) & m1_subset_1(sK1,u1_struct_0(sK0))) & l2_algstr_0(sK0) & v2_rlvect_1(sK0) & v3_algstr_1(sK0) & v5_algstr_0(sK0) & ~v2_struct_0(sK0) [skolemisation 27,43,42]
45. ! [X0] : (! [X1] : (! [X2] : (((k1_algstr_2(X0,X1) = X2 | k1_algstr_0(X0,X1,X2) != k4_struct_0(X0)) & (k1_algstr_0(X0,X1,X2) = k4_struct_0(X0) | k1_algstr_2(X0,X1) != X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0)) [nnf transformation 31]
54. ~v2_struct_0(sK0) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. v5_algstr_0(sK0) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. v3_algstr_1(sK0) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
57. v2_rlvect_1(sK0) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
58. l2_algstr_0(sK0) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
59. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 44] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
60. sK1 != k1_algstr_2(sK0,k1_algstr_2(sK0,sK1)) [cnf transformation 44] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
63. ~l2_algstr_0(X0) | l1_algstr_0(X0) [cnf transformation 29] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
64. k1_algstr_0(X0,X1,X2) = k4_struct_0(X0) | k1_algstr_2(X0,X1) != X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 45] {a:0,w:28,wCS:11200,sine:255,thAx:0,allAx:1,thDist:-1}
65. k1_algstr_0(X0,X1,X2) != k4_struct_0(X0) | k1_algstr_2(X0,X1) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 45] {a:0,w:28,wCS:11200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
67. m1_subset_1(k1_algstr_2(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 33] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. k3_rlvect_1(X0,X1,X2) = k1_algstr_0(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_algstr_0(X0) | ~v2_rlvect_1(X0) [cnf transformation 37] {a:0,w:21,wCS:8400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
70. k3_rlvect_1(X0,X1,X2) = k3_rlvect_1(X0,X2,X1) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_algstr_0(X0) | ~v2_rlvect_1(X0) [cnf transformation 39] {a:0,w:21,wCS:8400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
75. k4_struct_0(X0) = k1_algstr_0(X0,X1,k1_algstr_2(X0,X1)) | ~m1_subset_1(k1_algstr_2(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) [equality resolution 64] {a:0,w:27,wCS:10800,sine:255,thAx:0,allAx:1,thDist:-1}
77. 1 <=> sK1 = k1_algstr_2(sK0,k1_algstr_2(sK0,sK1)) [avatar definition]
79. sK1 != k1_algstr_2(sK0,k1_algstr_2(sK0,sK1)) <- (~1) [avatar component clause 77] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
80. ~1 [avatar split clause 60,77]
82. 2 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
84. m1_subset_1(sK1,u1_struct_0(sK0)) <- (2) [avatar component clause 82] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
85. 2 [avatar split clause 59,82]
87. 3 <=> l2_algstr_0(sK0) [avatar definition]
89. l2_algstr_0(sK0) <- (3) [avatar component clause 87] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
90. 3 [avatar split clause 58,87]
92. 4 <=> v2_rlvect_1(sK0) [avatar definition]
94. v2_rlvect_1(sK0) <- (4) [avatar component clause 92] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. 4 [avatar split clause 57,92]
97. 5 <=> v3_algstr_1(sK0) [avatar definition]
99. v3_algstr_1(sK0) <- (5) [avatar component clause 97] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. 5 [avatar split clause 56,97]
102. 6 <=> v5_algstr_0(sK0) [avatar definition]
104. v5_algstr_0(sK0) <- (6) [avatar component clause 102] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 6 [avatar split clause 55,102]
107. 7 <=> v2_struct_0(sK0) [avatar definition]
109. ~v2_struct_0(sK0) <- (~7) [avatar component clause 107] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. ~7 [avatar split clause 54,107]
156. k4_struct_0(X0) = k1_algstr_0(X0,X1,k1_algstr_2(X0,X1)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) [subsumption resolution 75,67] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
157. k3_rlvect_1(X2,X3,k1_algstr_2(X2,X3)) = k4_struct_0(X2) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l2_algstr_0(X2) | ~v3_algstr_1(X2) | ~v5_algstr_0(X2) | v2_struct_0(X2) | ~m1_subset_1(k1_algstr_2(X2,X3),u1_struct_0(X2)) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l1_algstr_0(X2) | ~v2_rlvect_1(X2) [superposition 156,69] {a:1,w:35,wCS:14000,sine:1,thAx:0,allAx:2,thDist:-2}
162. k3_rlvect_1(X2,X3,k1_algstr_2(X2,X3)) = k4_struct_0(X2) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l2_algstr_0(X2) | ~v3_algstr_1(X2) | ~v5_algstr_0(X2) | v2_struct_0(X2) | ~m1_subset_1(k1_algstr_2(X2,X3),u1_struct_0(X2)) | ~l1_algstr_0(X2) | ~v2_rlvect_1(X2) [duplicate literal removal 157] {a:1,w:31,wCS:12400,sine:1,thAx:0,allAx:2,thDist:-2}
163. k3_rlvect_1(X2,X3,k1_algstr_2(X2,X3)) = k4_struct_0(X2) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l2_algstr_0(X2) | ~v3_algstr_1(X2) | ~v5_algstr_0(X2) | v2_struct_0(X2) | ~m1_subset_1(k1_algstr_2(X2,X3),u1_struct_0(X2)) | ~v2_rlvect_1(X2) [subsumption resolution 162,63] {a:1,w:29,wCS:11600,sine:1,thAx:0,allAx:2,thDist:-2}
164. k3_rlvect_1(X2,X3,k1_algstr_2(X2,X3)) = k4_struct_0(X2) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l2_algstr_0(X2) | ~v3_algstr_1(X2) | ~v5_algstr_0(X2) | v2_struct_0(X2) | ~v2_rlvect_1(X2) [subsumption resolution 163,67] {a:1,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
169. k3_rlvect_1(X0,X1,X2) != k4_struct_0(X0) | k1_algstr_2(X0,X1) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_algstr_0(X0) | ~v2_rlvect_1(X0) [superposition 65,69] {a:1,w:40,wCS:16000,sine:255,thAx:0,allAx:2,thDist:-2}
171. k3_rlvect_1(X0,X1,X2) != k4_struct_0(X0) | k1_algstr_2(X0,X1) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) | ~l1_algstr_0(X0) | ~v2_rlvect_1(X0) [duplicate literal removal 169] {a:1,w:32,wCS:12800,sine:255,thAx:0,allAx:2,thDist:-2}
172. k3_rlvect_1(X0,X1,X2) != k4_struct_0(X0) | k1_algstr_2(X0,X1) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) | ~v2_rlvect_1(X0) [subsumption resolution 171,63] {a:1,w:30,wCS:12000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
175. k3_rlvect_1(X0,X2,X1) != k4_struct_0(X0) | k1_algstr_2(X0,X1) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) | ~v2_rlvect_1(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l1_algstr_0(X0) | ~v2_rlvect_1(X0) [superposition 172,70] {a:2,w:42,wCS:16800,sine:255,thAx:0,allAx:3,thDist:-3}
179. k3_rlvect_1(X0,X2,X1) != k4_struct_0(X0) | k1_algstr_2(X0,X1) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) | ~v2_rlvect_1(X0) | ~l1_algstr_0(X0) [duplicate literal removal 175] {a:2,w:32,wCS:12800,sine:255,thAx:0,allAx:3,thDist:-3}
180. k3_rlvect_1(X0,X2,X1) != k4_struct_0(X0) | k1_algstr_2(X0,X1) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v3_algstr_1(X0) | ~v5_algstr_0(X0) | v2_struct_0(X0) | ~v2_rlvect_1(X0) [subsumption resolution 179,63] {a:2,w:30,wCS:12000,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
184. k4_struct_0(X6) != k4_struct_0(X6) | k1_algstr_2(X6,k1_algstr_2(X6,X7)) = X7 | ~m1_subset_1(X7,u1_struct_0(X6)) | ~m1_subset_1(k1_algstr_2(X6,X7),u1_struct_0(X6)) | ~l2_algstr_0(X6) | ~v3_algstr_1(X6) | ~v5_algstr_0(X6) | v2_struct_0(X6) | ~v2_rlvect_1(X6) | ~m1_subset_1(X7,u1_struct_0(X6)) | ~l2_algstr_0(X6) | ~v3_algstr_1(X6) | ~v5_algstr_0(X6) | v2_struct_0(X6) | ~v2_rlvect_1(X6) [superposition 180,164] {a:3,w:46,wCS:18400,sine:1,thAx:0,allAx:5,thDist:-5}
185. k4_struct_0(X6) != k4_struct_0(X6) | k1_algstr_2(X6,k1_algstr_2(X6,X7)) = X7 | ~m1_subset_1(X7,u1_struct_0(X6)) | ~m1_subset_1(k1_algstr_2(X6,X7),u1_struct_0(X6)) | ~l2_algstr_0(X6) | ~v3_algstr_1(X6) | ~v5_algstr_0(X6) | v2_struct_0(X6) | ~v2_rlvect_1(X6) [duplicate literal removal 184] {a:3,w:32,wCS:12800,sine:1,thAx:0,allAx:5,thDist:-5}
186. k1_algstr_2(X6,k1_algstr_2(X6,X7)) = X7 | ~m1_subset_1(X7,u1_struct_0(X6)) | ~m1_subset_1(k1_algstr_2(X6,X7),u1_struct_0(X6)) | ~l2_algstr_0(X6) | ~v3_algstr_1(X6) | ~v5_algstr_0(X6) | v2_struct_0(X6) | ~v2_rlvect_1(X6) [trivial inequality removal 185] {a:3,w:27,wCS:10800,sine:1,thAx:0,allAx:5,thDist:-5}
189. k1_algstr_2(X6,k1_algstr_2(X6,X7)) = X7 | ~m1_subset_1(X7,u1_struct_0(X6)) | ~l2_algstr_0(X6) | ~v3_algstr_1(X6) | ~v5_algstr_0(X6) | v2_struct_0(X6) | ~v2_rlvect_1(X6) [subsumption resolution 186,67] {a:3,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
191. sK1 != sK1 | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l2_algstr_0(sK0) | ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) | ~v2_rlvect_1(sK0) <- (~1) [superposition 79,189] {a:4,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
198. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l2_algstr_0(sK0) | ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) | ~v2_rlvect_1(sK0) <- (~1) [trivial inequality removal 191] {a:4,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
200. ~l2_algstr_0(sK0) | ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) | ~v2_rlvect_1(sK0) <- (~1, 2) [subsumption resolution 198,84] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
201. ~v3_algstr_1(sK0) | ~v5_algstr_0(sK0) | v2_struct_0(sK0) | ~v2_rlvect_1(sK0) <- (~1, 2, 3) [subsumption resolution 200,89] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
202. ~v5_algstr_0(sK0) | v2_struct_0(sK0) | ~v2_rlvect_1(sK0) <- (~1, 2, 3, 5) [subsumption resolution 201,99] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
203. v2_struct_0(sK0) | ~v2_rlvect_1(sK0) <- (~1, 2, 3, 5, 6) [subsumption resolution 202,104] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
204. ~v2_rlvect_1(sK0) <- (~1, 2, 3, 5, 6, ~7) [subsumption resolution 203,109] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
205. $false <- (~1, 2, 3, 4, 5, 6, ~7) [subsumption resolution 204,94] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
206. 1 | ~2 | ~3 | ~4 | ~5 | ~6 | 7 [avatar contradiction clause 205]
209. $false [avatar sat refutation 80,85,90,95,100,105,110,206] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:93,thDist:-93}
% SZS output end Proof for algstr_2__t3_algstr_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 20
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 94
% Active clauses: 30
% Passive clauses: 33
% Final active clauses: 30
% Final passive clauses: 1
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Duplicate literals: 49
% Trivial inequalities: 2
% Fw subsumption resolutions: 19
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% Forward superposition: 7
% Backward superposition: 10
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 10
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

