i: [1,87,0,3,1,1,0,0] t47_setfam_1
a: [3,122,0,3,1,87]
s: 122
i: [1,86,0,3,1,1,0,0] t47_setfam_1
a: [3,127,0,3,1,86]
s: 127
i: [1,85,0,3,1,1,0,0] t47_setfam_1
a: [3,132,0,3,1,85]
s: 132
i: [1,84,0,5,1,1,0,0] t47_setfam_1
a: [3,137,0,5,1,84]
s: 137
i: [1,83,0,5,1,1,0,0] t47_setfam_1
a: [3,142,0,5,1,83]
s: 142
i: [1,88,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,147,0,2,1,88]
s: 147
i: [1,89,0,2,1,0,0,255] fc1_xboole_0
a: [3,152,0,2,1,89]
s: 152
i: [1,117,0,2,1,0,0,2] rc2_xboole_0
a: [3,162,0,2,1,117]
s: 162
i: [1,118,0,2,1,0,0,2] rc1_xboole_0
a: [3,167,0,2,1,118]
s: 167
i: [1,91,0,3,1,0,0,1] fc1_subset_1
s: 91
i: [1,112,0,3,1,0,0,1] rc2_subset_1
s: 112
i: [1,115,0,3,1,0,0,1] reflexivity_r1_tarski
s: 115
i: [1,90,0,3,1,0,0,255] d2_xboole_0
a: [3,157,0,3,1,90]
s: 157
r: 147
i: [1,110,0,4,1,0,0,1] existence_m1_subset_1
s: 110
i: [1,114,0,4,1,0,0,255] rc3_subset_1
s: 114
i: [1,95,0,5,2,0,0,1] rc1_subset_1
s: 95
i: [1,98,0,5,2,0,0,255] t6_boole
s: 98
i: [1,107,0,5,2,0,0,2] t7_boole
s: 107
i: [1,111,0,5,1,0,0,1] rc2_subset_1
s: 111
i: [1,113,0,5,1,0,0,255] rc3_subset_1
s: 113
i: [1,97,0,6,2,0,0,255] rc4_subset_1
s: 97
i: [1,101,0,6,2,0,0,2] t1_subset
s: 101
i: [1,116,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 116
i: [1,94,0,7,2,0,0,1] rc1_subset_1
s: 94
i: [1,96,0,7,2,0,0,255] rc4_subset_1
s: 96
i: [1,105,0,7,2,0,0,1] t3_subset
s: 105
i: [1,106,0,7,3,0,0,255] t8_boole
s: 106
i: [1,99,0,8,3,0,0,1] cc1_subset_1
s: 99
d: [2,177,1,5,2,0,40,99,113]
s: 177
i: [1,102,0,8,3,0,0,2] t2_subset
s: 102
i: [1,100,0,9,3,0,0,255] cc4_subset_1
s: 100
d: [2,180,1,9,2,1,40,102,137]
d: [2,188,1,5,1,1,21,180,91]
a: [3,192,1,5,1,188]
s: 192
i: [1,109,0,9,3,0,0,2] t5_subset
s: 109
d: [2,181,1,9,2,1,40,102,142]
d: [2,194,1,5,1,1,21,181,91]
a: [3,198,1,5,1,194]
s: 198
i: [1,108,0,10,3,0,0,2] t4_subset
s: 108
i: [1,92,0,11,4,0,0,255] cc3_subset_1
s: 92
d: [2,245,1,9,3,1,40,92,137]
d: [2,255,1,6,2,1,21,245,91]
a: [3,263,1,2,1,255]
s: 263
i: [1,93,0,11,4,0,0,255] cc2_subset_1
s: 93
d: [2,246,1,9,3,1,40,92,142]
d: [2,265,1,6,2,1,21,246,91]
a: [3,273,1,2,1,265]
s: 273
i: [1,103,0,11,3,0,0,1] d12_setfam_1
s: 103
i: [1,104,0,11,3,0,0,1] d12_setfam_1
s: 104
d: [2,278,1,6,2,1,40,103,142]
d: [2,285,1,3,1,2,21,278,132]
a: [3,289,1,3,1,285]
s: 289
d: [2,293,1,6,2,1,40,104,137]
d: [2,301,1,3,1,2,21,293,122]
a: [3,305,1,3,1,301]
s: 305
d: [2,309,2,3,1,1,40,305,116]
a: [3,315,2,3,1,309]
s: 315
d: [2,237,1,7,2,1,40,108,137]
s: 237
d: [2,238,1,7,2,1,40,108,142]
s: 238
d: [2,322,2,10,4,1,40,237,92]
a: [3,330,2,2,1,322]
s: 330
d: [2,184,1,8,2,0,40,102,111]
d: [2,202,1,5,1,0,21,184,91]
s: 202
d: [2,310,2,3,1,1,40,305,101]
a: [3,320,2,3,1,310]
s: 320
d: [2,185,1,8,2,0,40,102,113]
d: [2,203,1,5,1,0,21,185,91]
s: 203
d: [2,250,1,9,3,0,40,92,113]
d: [2,276,1,5,2,0,21,250,114]
s: 276
d: [2,213,2,5,1,1,40,192,116]
a: [3,219,2,5,1,213]
s: 219
d: [2,179,1,6,2,0,40,99,110]
s: 179
d: [2,229,2,5,1,1,40,198,116]
a: [3,235,2,5,1,229]
s: 235
d: [2,187,1,6,2,0,40,102,110]
s: 187
d: [2,209,1,6,2,0,40,100,111]
s: 209
d: [2,327,2,10,3,1,40,237,102]
d: [2,348,2,7,2,1,21,327,91]
s: 348
d: [2,225,1,6,2,0,40,109,111]
s: 225
d: [2,354,2,10,3,1,40,238,102]
d: [2,371,2,7,2,1,21,354,91]
s: 371
d: [2,226,1,6,2,0,40,109,113]
s: 226
d: [2,249,1,9,3,0,40,92,111]
d: [2,275,1,6,2,0,21,249,112]
s: 275
d: [2,383,3,7,2,1,40,348,116]
s: 383
d: [2,170,1,7,3,0,44,98,98]
s: 170
d: [2,392,4,8,2,2,40,383,371]
a: [3,402,4,4,1,392]
s: 402
d: [2,178,1,7,3,0,40,99,105]
s: 178
d: [2,182,1,10,3,0,40,102,94]
d: [2,200,1,7,2,0,21,182,91]
s: 200
d: [2,387,3,7,2,1,40,371,116]
s: 387
d: [2,183,1,10,3,0,40,102,96]
d: [2,201,1,7,2,0,21,183,91]
s: 201
a: [3,398,4,4,1,392]
s: 398
d: [2,186,1,10,3,0,40,102,105]
d: [2,204,1,7,2,0,21,186,91]
s: 204
d: [2,212,1,7,2,0,40,100,110]
s: 212
a: [3,334,2,8,3,322]
s: 334
d: [2,228,1,7,2,0,40,109,110]
s: 228
d: [2,423,3,9,3,1,40,334,187]
d: [2,424,3,7,2,2,21,423,263]
a: [3,428,3,3,1,424]
s: 428
d: [2,241,1,7,2,0,40,108,111]
s: 241
d: [2,242,1,7,2,0,40,108,113]
s: 242
d: [2,349,2,10,4,1,40,238,92]
a: [3,357,2,8,3,349]
s: 357
d: [2,211,1,8,3,0,40,100,105]
s: 211
d: [2,437,3,9,3,1,40,357,187]
d: [2,438,3,7,2,2,21,437,273]
a: [3,442,3,3,1,438]
s: 442
d: [2,227,1,8,3,0,40,109,105]
s: 227
d: [2,244,1,8,2,0,40,108,110]
s: 244
d: [2,459,2,5,2,1,40,227,305]
s: 459
d: [2,239,1,9,3,0,40,108,94]
s: 239
d: [2,457,3,9,3,1,40,227,371]
a: [3,466,3,6,2,457]
s: 466
d: [2,240,1,9,3,0,40,108,96]
s: 240
d: [2,243,1,9,3,0,40,108,105]
s: 243
d: [2,488,2,6,2,1,40,243,305]
s: 488
d: [2,251,1,10,4,0,40,92,105]
s: 251
d: [2,490,3,3,1,2,40,488,127]
a: [3,495,3,3,1,490]
s: 495
d: [2,499,4,0,0,3,52,273,289,495,102]
e: 499
d: [2,500,4,5,2,1,40,495,102]
d: [2,502,4,3,1,2,21,500,273]
d: [2,503,4,0,0,3,21,502,289]
e: 503
r: 273
d: [2,505,4,2,1,2,21,500,289]
d: [2,473,3,2,1,2,40,459,127]
f: 473,505,490,301,285,83,84,85,86,87
% Refutation found. Thanks to Tanya!
% SZS status Theorem for setfam_1__t47_setfam_1
% SZS output start Proof for setfam_1__t47_setfam_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) => ((r1_tarski(X2,X1) & v3_setfam_1(X1,X0)) => v3_setfam_1(X2,X0)))) [input t47_setfam_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) => ((r1_tarski(X2,X1) & v3_setfam_1(X1,X0)) => v3_setfam_1(X2,X0)))) [negated conjecture 1]
8. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => (v3_setfam_1(X1,X0) <=> ~r2_hidden(X0,X1))) [input d12_setfam_1]
25. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
26. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
27. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
28. ! [X0] : ! [X1] : ! [X2] : ~(v1_xboole_0(X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) [input t5_subset]
32. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) => ((r1_tarski(X2,X1) & v3_setfam_1(X1,X0)) => v3_setfam_1(X2,X0)))) [flattening 2]
34. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 25]
35. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => (v3_setfam_1(X1,X0) <=> ~r2_hidden(X0,X1))) [flattening 8]
36. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 26]
39. ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 27]
40. ! [X0,X1,X2] : ~(v1_xboole_0(X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) [flattening 28]
42. ! [X0,X1] : (r1_tarski(X0,X1) => m1_subset_1(X0,k1_zfmisc_1(X1))) [unused predicate definition removal 36]
43. ? [X0,X1] : (? [X2] : ((~v3_setfam_1(X2,X0) & (r1_tarski(X2,X1) & v3_setfam_1(X1,X0))) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0)))) & m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [ennf transformation 32]
44. ? [X0,X1] : (? [X2] : (~v3_setfam_1(X2,X0) & r1_tarski(X2,X1) & v3_setfam_1(X1,X0) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0)))) & m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [flattening 43]
55. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 34]
56. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 55]
57. ! [X0,X1] : ((v3_setfam_1(X1,X0) <=> ~r2_hidden(X0,X1)) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [ennf transformation 35]
58. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 42]
61. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))) [ennf transformation 39]
62. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [flattening 61]
63. ! [X0,X1,X2] : (~v1_xboole_0(X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [ennf transformation 40]
65. ? [X0,X1] : (? [X2] : (~v3_setfam_1(X2,X0) & r1_tarski(X2,X1) & v3_setfam_1(X1,X0) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0)))) & m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) => (? [X2] : (~v3_setfam_1(X2,sK0) & r1_tarski(X2,sK1) & v3_setfam_1(sK1,sK0) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(sK0)))) & m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0)))) [choice axiom]
66. ? [X2] : (~v3_setfam_1(X2,sK0) & r1_tarski(X2,sK1) & v3_setfam_1(sK1,sK0) & m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(sK0)))) => (~v3_setfam_1(sK2,sK0) & r1_tarski(sK2,sK1) & v3_setfam_1(sK1,sK0) & m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(sK0)))) [choice axiom]
67. (~v3_setfam_1(sK2,sK0) & r1_tarski(sK2,sK1) & v3_setfam_1(sK1,sK0) & m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(sK0)))) & m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) [skolemisation 44,66,65]
72. ! [X0,X1] : (((v3_setfam_1(X1,X0) | r2_hidden(X0,X1)) & (~r2_hidden(X0,X1) | ~v3_setfam_1(X1,X0))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [nnf transformation 57]
83. m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) [cnf transformation 67] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
84. m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(sK0))) [cnf transformation 67] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
85. v3_setfam_1(sK1,sK0) [cnf transformation 67] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
86. r1_tarski(sK2,sK1) [cnf transformation 67] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
87. ~v3_setfam_1(sK2,sK0) [cnf transformation 67] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
102. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 56] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
103. ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v3_setfam_1(X1,X0) | ~r2_hidden(X0,X1) [cnf transformation 72] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
104. ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | r2_hidden(X0,X1) | v3_setfam_1(X1,X0) [cnf transformation 72] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
105. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 58] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
108. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1) [cnf transformation 62] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
109. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~v1_xboole_0(X2) | ~r2_hidden(X0,X1) [cnf transformation 63] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
120. 1 <=> v3_setfam_1(sK2,sK0) [avatar definition]
122. ~v3_setfam_1(sK2,sK0) <- (~1) [avatar component clause 120] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
123. ~1 [avatar split clause 87,120]
125. 2 <=> r1_tarski(sK2,sK1) [avatar definition]
127. r1_tarski(sK2,sK1) <- (2) [avatar component clause 125] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
128. 2 [avatar split clause 86,125]
130. 3 <=> v3_setfam_1(sK1,sK0) [avatar definition]
132. v3_setfam_1(sK1,sK0) <- (3) [avatar component clause 130] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
133. 3 [avatar split clause 85,130]
135. 4 <=> m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(sK0))) [avatar definition]
137. m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(sK0))) <- (4) [avatar component clause 135] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
138. 4 [avatar split clause 84,135]
140. 5 <=> m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) [avatar definition]
142. m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) <- (5) [avatar component clause 140] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
143. 5 [avatar split clause 83,140]
227. ~r2_hidden(X11,X12) | ~v1_xboole_0(X10) | ~r1_tarski(X12,X10) [resolution 109,105] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
243. ~r2_hidden(X10,X12) | m1_subset_1(X10,X11) | ~r1_tarski(X12,X11) [resolution 108,105] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
271. 18 <=> v1_xboole_0(sK1) [avatar definition]
278. ~v3_setfam_1(sK1,sK0) | ~r2_hidden(sK0,sK1) <- (5) [resolution 103,142] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
285. ~r2_hidden(sK0,sK1) <- (3, 5) [subsumption resolution 278,132] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
287. 19 <=> r2_hidden(sK0,sK1) [avatar definition]
289. ~r2_hidden(sK0,sK1) <- (~19) [avatar component clause 287] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
290. ~19 | ~3 | ~5 [avatar split clause 285,140,130,287]
293. r2_hidden(sK0,sK2) | v3_setfam_1(sK2,sK0) <- (4) [resolution 104,137] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
301. r2_hidden(sK0,sK2) <- (~1, 4) [subsumption resolution 293,122] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
303. 20 <=> r2_hidden(sK0,sK2) [avatar definition]
305. r2_hidden(sK0,sK2) <- (20) [avatar component clause 303] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
306. 20 | 1 | ~4 [avatar split clause 301,135,120,303]
459. ~r1_tarski(sK2,X17) | ~v1_xboole_0(X17) <- (20) [resolution 227,305] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
473. ~v1_xboole_0(sK1) <- (2, 20) [resolution 459,127] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
488. ~r1_tarski(sK2,X17) | m1_subset_1(sK0,X17) <- (20) [resolution 243,305] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
490. m1_subset_1(sK0,sK1) <- (2, 20) [resolution 488,127] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
493. 41 <=> m1_subset_1(sK0,sK1) [avatar definition]
495. m1_subset_1(sK0,sK1) <- (41) [avatar component clause 493] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
496. 41 | ~2 | ~20 [avatar split clause 490,303,125,493]
500. v1_xboole_0(sK1) | r2_hidden(sK0,sK1) <- (41) [resolution 495,102] {a:4,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
505. v1_xboole_0(sK1) <- (~19, 41) [subsumption resolution 500,289] {a:4,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:10,thDist:-10}
506. 18 | 19 | ~41 [avatar split clause 505,493,287,271]
507. ~18 | ~2 | ~20 [avatar split clause 473,303,125,271]
510. $false [avatar sat refutation 123,128,133,138,143,290,306,496,506,507] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:322,thDist:-322}
% SZS output end Proof for setfam_1__t47_setfam_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 30
% >>> Preprocessing
% Introduced skolems: 10
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 268
% Active clauses: 87
% Passive clauses: 135
% Final active clauses: 85
% Final passive clauses: 46
% Main loop iterations started: 89
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 26
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 50
% 
% >>> Generating Inferences
% Binary resolution: 163
% Unit resulting resolution: 1
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 15
% Split components: 30
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 43
% SAT solver unit clauses: 10
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.002 s
% other: 0.003 s ( own 0.001 s ) 

