i: [1,99,0,2,1,1,0,0] t6_aff_1
a: [3,182,0,2,1,99]
s: 182
i: [1,98,0,2,1,1,0,0] t6_aff_1
a: [3,187,0,2,1,98]
s: 187
i: [1,97,0,2,1,1,0,0] t6_aff_1
a: [3,192,0,2,1,97]
s: 192
i: [1,102,0,4,1,1,0,0] t6_aff_1
a: [3,167,0,4,1,102]
s: 167
i: [1,101,0,4,1,1,0,0] t6_aff_1
a: [3,172,0,4,1,101]
s: 172
i: [1,100,0,4,1,1,0,0] t6_aff_1
a: [3,177,0,4,1,100]
s: 177
i: [1,104,0,25,5,1,0,0] t6_aff_1
a: [3,141,0,5,1,104]
s: 141
i: [1,103,0,5,1,1,0,0] t6_aff_1
a: [3,162,0,5,1,103]
s: 162
i: [1,105,0,2,1,0,0,255] spc1_boole
a: [3,197,0,2,1,105]
s: 197
i: [1,136,0,2,1,0,0,1] existence_l1_analoaf
a: [3,202,0,2,1,136]
s: 202
i: [1,137,0,2,1,0,0,2] existence_l1_struct_0
a: [3,207,0,2,1,137]
s: 207
i: [1,106,0,4,2,0,0,1] dt_l1_analoaf
s: 106
i: [1,134,0,4,1,0,0,1] existence_m1_subset_1
s: 134
d: [2,209,1,2,1,1,40,106,182]
a: [3,214,1,2,1,209]
s: 214
i: [1,107,0,5,2,0,0,255] t6_boole
s: 107
i: [1,133,0,5,2,0,0,255] t7_boole
s: 133
d: [2,210,1,2,1,1,40,106,202]
a: [3,219,1,2,1,210]
s: 219
i: [1,108,0,6,3,0,0,255] cc7_struct_0
s: 108
d: [2,222,1,4,2,1,40,108,207]
a: [3,228,1,2,1,222]
s: 228
i: [1,109,0,6,3,0,0,255] cc5_struct_0
s: 109
i: [1,111,0,6,3,0,0,1] cc2_struct_0
s: 111
d: [2,224,2,4,2,1,40,108,219]
a: [3,237,2,2,1,224]
s: 237
i: [1,112,0,6,3,0,0,1] cc17_struct_0
s: 112
i: [1,113,0,6,3,0,0,1] cc1_struct_0
s: 113
i: [1,116,0,6,3,0,0,255] cc4_struct_0
s: 116
i: [1,117,0,6,3,0,0,255] cc6_struct_0
s: 117
i: [1,130,0,6,2,0,0,255] t1_subset
s: 130
i: [1,135,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 135
i: [1,114,0,7,3,0,0,255] cc8_struct_0
s: 114
i: [1,118,0,7,3,0,0,255] cc9_struct_0
s: 118
i: [1,119,0,7,3,0,0,255] cc11_struct_0
s: 119
i: [1,120,0,7,3,0,0,255] cc11_struct_0
s: 120
i: [1,122,0,7,3,0,0,255] fc2_struct_0
s: 122
i: [1,123,0,7,3,0,0,255] fc9_struct_0
s: 123
i: [1,126,0,7,3,0,0,255] fc6_struct_0
s: 126
i: [1,127,0,7,3,0,0,255] fc1_struct_0
s: 127
i: [1,128,0,7,3,0,0,255] fc8_struct_0
s: 128
i: [1,129,0,7,3,0,0,255] fc7_struct_0
s: 129
i: [1,132,0,7,3,0,0,255] t8_boole
s: 132
i: [1,131,0,8,3,0,0,255] t2_subset
s: 131
i: [1,121,0,9,4,0,0,255] cc10_struct_0
s: 121
d: [2,253,1,7,2,1,40,131,167]
a: [3,264,1,3,1,253]
s: 264
r: 264
i: [1,124,0,28,8,0,0,1] l10_aff_1
s: 124
d: [2,287,1,0,0,8,52,192,187,182,177,172,162,141,167,124]
e: 287
d: [2,288,1,23,7,1,40,124,162]
d: [2,290,1,18,6,2,21,288,141]
d: [2,291,1,14,5,3,21,290,167]
d: [2,292,1,10,4,4,21,291,172]
d: [2,293,1,6,3,5,21,292,177]
d: [2,294,1,4,2,6,21,293,182]
d: [2,295,1,2,1,7,21,294,187]
d: [2,296,1,0,0,8,21,295,192]
e: 296
r: 141
a: [3,263,1,3,1,253]
s: 263
a: [3,140,0,5,1,104]
s: 140
a: [3,145,0,5,1,104]
s: 145
d: [2,280,2,4,2,1,40,264,122]
d: [2,281,2,2,1,2,21,280,214]
a: [3,284,2,2,1,281]
s: 284
i: [1,125,0,28,8,0,0,1] l10_aff_1
s: 125
d: [2,309,1,0,0,8,52,192,187,182,177,172,162,145,167,125]
e: 309
d: [2,310,1,23,7,1,40,125,162]
d: [2,313,1,18,6,2,21,310,145]
d: [2,314,1,14,5,3,21,313,167]
d: [2,315,1,10,4,4,21,314,172]
d: [2,316,1,6,3,5,21,315,177]
d: [2,317,1,4,2,6,21,316,182]
d: [2,318,1,2,1,7,21,317,187]
d: [2,319,1,0,0,8,21,318,192]
e: 319
r: 145
a: [3,260,1,4,1,253]
s: 260
a: [3,144,0,5,1,104]
s: 144
a: [3,148,0,5,1,104]
s: 148
d: [2,298,1,19,6,2,21,288,167]
d: [2,299,1,15,5,3,21,298,172]
d: [2,300,1,11,4,4,21,299,177]
d: [2,301,1,9,3,5,21,300,182]
d: [2,302,1,7,2,6,21,301,187]
d: [2,303,1,5,1,7,21,302,192]
d: [2,327,1,19,6,2,21,310,167]
d: [2,328,1,15,5,3,21,327,172]
d: [2,329,1,11,4,4,21,328,177]
d: [2,330,1,9,3,5,21,329,182]
d: [2,331,1,7,2,6,21,330,187]
d: [2,332,1,5,1,7,21,331,192]
d: [2,311,1,23,7,1,40,125,140]
d: [2,321,1,19,6,2,21,311,172]
d: [2,322,1,15,5,3,21,321,167]
d: [2,323,1,11,4,4,21,322,177]
d: [2,324,1,9,3,5,21,323,182]
d: [2,325,1,7,2,6,21,324,187]
d: [2,326,1,5,1,7,21,325,192]
d: [2,344,1,23,7,1,40,144,124]
d: [2,345,1,19,6,2,21,344,167]
d: [2,346,1,15,5,3,21,345,177]
d: [2,347,1,11,4,4,21,346,172]
d: [2,348,1,9,3,5,21,347,182]
d: [2,349,1,7,2,6,21,348,187]
d: [2,350,1,5,1,7,21,349,192]
d: [2,352,1,23,7,1,40,148,125]
d: [2,354,1,19,6,2,21,352,177]
d: [2,355,1,15,5,3,21,354,167]
d: [2,356,1,11,4,4,21,355,172]
d: [2,357,1,9,3,5,21,356,182]
d: [2,358,1,7,2,6,21,357,187]
d: [2,359,1,5,1,7,21,358,192]
f: 359,350,326,332,303,97,98,99,100,101,102,103,104
% Refutation found. Thanks to Tanya!
% SZS status Theorem for aff_1__t6_aff_1
% SZS output start Proof for aff_1__t6_aff_1
1. ! [X0] : ((l1_analoaf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_aff_1(X0,X1,X2,X3) => (r1_aff_1(X0,X3,X2,X1) & r1_aff_1(X0,X3,X1,X2) & r1_aff_1(X0,X2,X3,X1) & r1_aff_1(X0,X2,X1,X3) & r1_aff_1(X0,X1,X3,X2))))))) [input t6_aff_1]
2. ~! [X0] : ((l1_analoaf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_aff_1(X0,X1,X2,X3) => (r1_aff_1(X0,X3,X2,X1) & r1_aff_1(X0,X3,X1,X2) & r1_aff_1(X0,X2,X3,X1) & r1_aff_1(X0,X2,X1,X3) & r1_aff_1(X0,X1,X3,X2))))))) [negated conjecture 1]
29. ! [X0] : ((l1_analoaf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_aff_1(X0,X1,X2,X3) => (r1_aff_1(X0,X2,X1,X3) & r1_aff_1(X0,X1,X3,X2))))))) [input l10_aff_1]
40. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~r1_aff_1(X0,X3,X2,X1) | ~r1_aff_1(X0,X3,X1,X2) | ~r1_aff_1(X0,X2,X3,X1) | ~r1_aff_1(X0,X2,X1,X3) | ~r1_aff_1(X0,X1,X3,X2)) & r1_aff_1(X0,X1,X2,X3)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v1_diraf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
41. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_aff_1(X0,X3,X2,X1) | ~r1_aff_1(X0,X3,X1,X2) | ~r1_aff_1(X0,X2,X3,X1) | ~r1_aff_1(X0,X2,X1,X3) | ~r1_aff_1(X0,X1,X3,X2)) & r1_aff_1(X0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) [flattening 40]
70. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r1_aff_1(X0,X2,X1,X3) & r1_aff_1(X0,X1,X3,X2)) | ~r1_aff_1(X0,X1,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_diraf(X0) | v7_struct_0(X0))) [ennf transformation 29]
71. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_aff_1(X0,X2,X1,X3) & r1_aff_1(X0,X1,X3,X2)) | ~r1_aff_1(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_diraf(X0) | v7_struct_0(X0)) [flattening 70]
86. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_aff_1(X0,X3,X2,X1) | ~r1_aff_1(X0,X3,X1,X2) | ~r1_aff_1(X0,X2,X3,X1) | ~r1_aff_1(X0,X2,X1,X3) | ~r1_aff_1(X0,X1,X3,X2)) & r1_aff_1(X0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_diraf(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : ((~r1_aff_1(sK0,X3,X2,X1) | ~r1_aff_1(sK0,X3,X1,X2) | ~r1_aff_1(sK0,X2,X3,X1) | ~r1_aff_1(sK0,X2,X1,X3) | ~r1_aff_1(sK0,X1,X3,X2)) & r1_aff_1(sK0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
87. ? [X1] : (? [X2] : (? [X3] : ((~r1_aff_1(sK0,X3,X2,X1) | ~r1_aff_1(sK0,X3,X1,X2) | ~r1_aff_1(sK0,X2,X3,X1) | ~r1_aff_1(sK0,X2,X1,X3) | ~r1_aff_1(sK0,X1,X3,X2)) & r1_aff_1(sK0,X1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : ((~r1_aff_1(sK0,X3,X2,sK1) | ~r1_aff_1(sK0,X3,sK1,X2) | ~r1_aff_1(sK0,X2,X3,sK1) | ~r1_aff_1(sK0,X2,sK1,X3) | ~r1_aff_1(sK0,sK1,X3,X2)) & r1_aff_1(sK0,sK1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
88. ? [X2] : (? [X3] : ((~r1_aff_1(sK0,X3,X2,sK1) | ~r1_aff_1(sK0,X3,sK1,X2) | ~r1_aff_1(sK0,X2,X3,sK1) | ~r1_aff_1(sK0,X2,sK1,X3) | ~r1_aff_1(sK0,sK1,X3,X2)) & r1_aff_1(sK0,sK1,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : ((~r1_aff_1(sK0,X3,sK2,sK1) | ~r1_aff_1(sK0,X3,sK1,sK2) | ~r1_aff_1(sK0,sK2,X3,sK1) | ~r1_aff_1(sK0,sK2,sK1,X3) | ~r1_aff_1(sK0,sK1,X3,sK2)) & r1_aff_1(sK0,sK1,sK2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
89. ? [X3] : ((~r1_aff_1(sK0,X3,sK2,sK1) | ~r1_aff_1(sK0,X3,sK1,sK2) | ~r1_aff_1(sK0,sK2,X3,sK1) | ~r1_aff_1(sK0,sK2,sK1,X3) | ~r1_aff_1(sK0,sK1,X3,sK2)) & r1_aff_1(sK0,sK1,sK2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) => ((~r1_aff_1(sK0,sK3,sK2,sK1) | ~r1_aff_1(sK0,sK3,sK1,sK2) | ~r1_aff_1(sK0,sK2,sK3,sK1) | ~r1_aff_1(sK0,sK2,sK1,sK3) | ~r1_aff_1(sK0,sK1,sK3,sK2)) & r1_aff_1(sK0,sK1,sK2,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
90. ((((~r1_aff_1(sK0,sK3,sK2,sK1) | ~r1_aff_1(sK0,sK3,sK1,sK2) | ~r1_aff_1(sK0,sK2,sK3,sK1) | ~r1_aff_1(sK0,sK2,sK1,sK3) | ~r1_aff_1(sK0,sK1,sK3,sK2)) & r1_aff_1(sK0,sK1,sK2,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_diraf(sK0) & ~v7_struct_0(sK0) [skolemisation 41,89,88,87,86]
97. ~v7_struct_0(sK0) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
98. v1_diraf(sK0) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
99. l1_analoaf(sK0) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
100. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 90] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
101. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 90] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
102. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 90] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
103. r1_aff_1(sK0,sK1,sK2,sK3) [cnf transformation 90] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
104. ~r1_aff_1(sK0,sK3,sK2,sK1) | ~r1_aff_1(sK0,sK3,sK1,sK2) | ~r1_aff_1(sK0,sK2,sK3,sK1) | ~r1_aff_1(sK0,sK2,sK1,sK3) | ~r1_aff_1(sK0,sK1,sK3,sK2) [cnf transformation 90] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
124. ~r1_aff_1(X0,X1,X2,X3) | r1_aff_1(X0,X1,X3,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 71] {a:0,w:28,wCS:11200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
125. ~r1_aff_1(X0,X1,X2,X3) | r1_aff_1(X0,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_diraf(X0) | v7_struct_0(X0) [cnf transformation 71] {a:0,w:28,wCS:11200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
139. 1 <=> r1_aff_1(sK0,sK1,sK3,sK2) [avatar definition]
140. r1_aff_1(sK0,sK1,sK3,sK2) <- (1) [avatar component clause 139] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
143. 2 <=> r1_aff_1(sK0,sK2,sK1,sK3) [avatar definition]
144. r1_aff_1(sK0,sK2,sK1,sK3) <- (2) [avatar component clause 143] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
147. 3 <=> r1_aff_1(sK0,sK2,sK3,sK1) [avatar definition]
148. r1_aff_1(sK0,sK2,sK3,sK1) <- (3) [avatar component clause 147] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
151. 4 <=> r1_aff_1(sK0,sK3,sK1,sK2) [avatar definition]
155. 5 <=> r1_aff_1(sK0,sK3,sK2,sK1) [avatar definition]
158. ~1 | ~2 | ~3 | ~4 | ~5 [avatar split clause 104,155,151,147,143,139]
160. 6 <=> r1_aff_1(sK0,sK1,sK2,sK3) [avatar definition]
162. r1_aff_1(sK0,sK1,sK2,sK3) <- (6) [avatar component clause 160] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
163. 6 [avatar split clause 103,160]
165. 7 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
167. m1_subset_1(sK3,u1_struct_0(sK0)) <- (7) [avatar component clause 165] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
168. 7 [avatar split clause 102,165]
170. 8 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
172. m1_subset_1(sK2,u1_struct_0(sK0)) <- (8) [avatar component clause 170] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
173. 8 [avatar split clause 101,170]
175. 9 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
177. m1_subset_1(sK1,u1_struct_0(sK0)) <- (9) [avatar component clause 175] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
178. 9 [avatar split clause 100,175]
180. 10 <=> l1_analoaf(sK0) [avatar definition]
182. l1_analoaf(sK0) <- (10) [avatar component clause 180] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
183. 10 [avatar split clause 99,180]
185. 11 <=> v1_diraf(sK0) [avatar definition]
187. v1_diraf(sK0) <- (11) [avatar component clause 185] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
188. 11 [avatar split clause 98,185]
190. 12 <=> v7_struct_0(sK0) [avatar definition]
192. ~v7_struct_0(sK0) <- (~12) [avatar component clause 190] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
193. ~12 [avatar split clause 97,190]
288. r1_aff_1(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6) [resolution 124,162] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
298. r1_aff_1(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6, 7) [subsumption resolution 288,167] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
299. r1_aff_1(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6, 7, 8) [subsumption resolution 298,172] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
300. r1_aff_1(sK0,sK1,sK3,sK2) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6, 7, 8, 9) [subsumption resolution 299,177] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
301. r1_aff_1(sK0,sK1,sK3,sK2) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6, 7, 8, 9, 10) [subsumption resolution 300,182] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
302. r1_aff_1(sK0,sK1,sK3,sK2) | v7_struct_0(sK0) <- (6, 7, 8, 9, 10, 11) [subsumption resolution 301,187] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
303. r1_aff_1(sK0,sK1,sK3,sK2) <- (6, 7, 8, 9, 10, 11, ~12) [subsumption resolution 302,192] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
304. 1 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | 12 [avatar split clause 303,190,185,180,175,170,165,160,139]
310. r1_aff_1(sK0,sK2,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6) [resolution 125,162] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
311. r1_aff_1(sK0,sK3,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1) [resolution 125,140] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
321. r1_aff_1(sK0,sK3,sK1,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1, 8) [subsumption resolution 311,172] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
322. r1_aff_1(sK0,sK3,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1, 7, 8) [subsumption resolution 321,167] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
323. r1_aff_1(sK0,sK3,sK1,sK2) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1, 7, 8, 9) [subsumption resolution 322,177] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
324. r1_aff_1(sK0,sK3,sK1,sK2) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (1, 7, 8, 9, 10) [subsumption resolution 323,182] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
325. r1_aff_1(sK0,sK3,sK1,sK2) | v7_struct_0(sK0) <- (1, 7, 8, 9, 10, 11) [subsumption resolution 324,187] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
326. r1_aff_1(sK0,sK3,sK1,sK2) <- (1, 7, 8, 9, 10, 11, ~12) [subsumption resolution 325,192] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
327. r1_aff_1(sK0,sK2,sK1,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6, 7) [subsumption resolution 310,167] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
328. r1_aff_1(sK0,sK2,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6, 7, 8) [subsumption resolution 327,172] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
329. r1_aff_1(sK0,sK2,sK1,sK3) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6, 7, 8, 9) [subsumption resolution 328,177] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
330. r1_aff_1(sK0,sK2,sK1,sK3) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (6, 7, 8, 9, 10) [subsumption resolution 329,182] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
331. r1_aff_1(sK0,sK2,sK1,sK3) | v7_struct_0(sK0) <- (6, 7, 8, 9, 10, 11) [subsumption resolution 330,187] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
332. r1_aff_1(sK0,sK2,sK1,sK3) <- (6, 7, 8, 9, 10, 11, ~12) [subsumption resolution 331,192] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
333. 2 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | 12 [avatar split clause 332,190,185,180,175,170,165,160,143]
334. 4 | ~1 | ~7 | ~8 | ~9 | ~10 | ~11 | 12 [avatar split clause 326,190,185,180,175,170,165,139,151]
344. r1_aff_1(sK0,sK2,sK3,sK1) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2) [resolution 144,124] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
345. r1_aff_1(sK0,sK2,sK3,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 7) [subsumption resolution 344,167] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
346. r1_aff_1(sK0,sK2,sK3,sK1) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 7, 9) [subsumption resolution 345,177] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
347. r1_aff_1(sK0,sK2,sK3,sK1) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 7, 8, 9) [subsumption resolution 346,172] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
348. r1_aff_1(sK0,sK2,sK3,sK1) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (2, 7, 8, 9, 10) [subsumption resolution 347,182] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
349. r1_aff_1(sK0,sK2,sK3,sK1) | v7_struct_0(sK0) <- (2, 7, 8, 9, 10, 11) [subsumption resolution 348,187] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
350. r1_aff_1(sK0,sK2,sK3,sK1) <- (2, 7, 8, 9, 10, 11, ~12) [subsumption resolution 349,192] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
351. 3 | ~2 | ~7 | ~8 | ~9 | ~10 | ~11 | 12 [avatar split clause 350,190,185,180,175,170,165,143,147]
352. r1_aff_1(sK0,sK3,sK2,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3) [resolution 148,125] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
354. r1_aff_1(sK0,sK3,sK2,sK1) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 9) [subsumption resolution 352,177] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
355. r1_aff_1(sK0,sK3,sK2,sK1) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 7, 9) [subsumption resolution 354,167] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
356. r1_aff_1(sK0,sK3,sK2,sK1) | ~l1_analoaf(sK0) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 7, 8, 9) [subsumption resolution 355,172] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
357. r1_aff_1(sK0,sK3,sK2,sK1) | ~v1_diraf(sK0) | v7_struct_0(sK0) <- (3, 7, 8, 9, 10) [subsumption resolution 356,182] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
358. r1_aff_1(sK0,sK3,sK2,sK1) | v7_struct_0(sK0) <- (3, 7, 8, 9, 10, 11) [subsumption resolution 357,187] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
359. r1_aff_1(sK0,sK3,sK2,sK1) <- (3, 7, 8, 9, 10, 11, ~12) [subsumption resolution 358,192] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
360. 5 | ~3 | ~7 | ~8 | ~9 | ~10 | ~11 | 12 [avatar split clause 359,190,185,180,175,170,165,147,155]
361. $false [avatar sat refutation 158,163,168,173,178,183,188,193,304,333,334,351,360] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:306,thDist:-306}
% SZS output end Proof for aff_1__t6_aff_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 34
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 41
% Generated clauses: 158
% Active clauses: 47
% Passive clauses: 61
% Final active clauses: 44
% Final passive clauses: 9
% Main loop iterations started: 51
% 
% >>> Simplifying Inferences
% Duplicate literals: 6
% Fw subsumption resolutions: 46
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 28
% Unit resulting resolution: 2
% Forward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 15
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 32
% SAT solver unit clauses: 10
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

