i: [1,112,0,2,1,1,0,0] l3_afvect01
a: [3,190,0,2,1,112]
s: 190
i: [1,111,0,2,1,1,0,0] l3_afvect01
a: [3,195,0,2,1,111]
s: 195
i: [1,110,0,2,1,1,0,0] l3_afvect01
a: [3,200,0,2,1,110]
s: 200
i: [1,116,0,4,1,1,0,0] l3_afvect01
a: [3,170,0,4,1,116]
s: 170
i: [1,115,0,4,1,1,0,0] l3_afvect01
a: [3,175,0,4,1,115]
s: 175
i: [1,114,0,4,1,1,0,0] l3_afvect01
a: [3,180,0,4,1,114]
s: 180
i: [1,113,0,4,1,1,0,0] l3_afvect01
a: [3,185,0,4,1,113]
s: 185
i: [1,118,0,6,1,1,0,0] l3_afvect01
a: [3,160,0,6,1,118]
s: 160
i: [1,117,0,6,1,1,0,0] l3_afvect01
a: [3,165,0,6,1,117]
s: 165
i: [1,119,0,2,1,0,0,255] spc1_boole
a: [3,205,0,2,1,119]
s: 205
i: [1,120,0,2,1,0,0,255] fc1_xboole_0
a: [3,210,0,2,1,120]
s: 210
i: [1,153,0,2,1,0,0,2] rc2_xboole_0
a: [3,215,0,2,1,153]
s: 215
i: [1,154,0,2,1,0,0,1] existence_l1_analoaf
a: [3,220,0,2,1,154]
s: 220
i: [1,155,0,2,1,0,0,2] existence_l1_struct_0
a: [3,225,0,2,1,155]
s: 225
i: [1,156,0,2,1,0,0,2] rc1_xboole_0
a: [3,230,0,2,1,156]
s: 230
i: [1,121,0,4,2,0,0,1] dt_l1_analoaf
s: 121
d: [2,232,1,2,1,1,40,121,190]
a: [3,237,1,2,1,232]
s: 237
i: [1,151,0,4,1,0,0,1] existence_m1_subset_1
s: 151
d: [2,233,1,2,1,1,40,121,220]
a: [3,242,1,2,1,233]
s: 242
i: [1,136,0,5,2,0,0,255] t6_boole
s: 136
i: [1,150,0,5,2,0,0,255] t7_boole
s: 150
i: [1,122,0,6,3,0,0,255] cc7_struct_0
s: 122
i: [1,123,0,6,3,0,0,255] cc5_struct_0
s: 123
d: [2,246,1,4,2,1,40,122,225]
a: [3,252,1,2,1,246]
s: 252
i: [1,125,0,6,3,0,0,1] cc17_struct_0
s: 125
i: [1,126,0,6,3,0,0,1] cc2_struct_0
s: 126
d: [2,248,2,4,2,1,40,122,242]
a: [3,261,2,2,1,248]
s: 261
i: [1,127,0,6,3,0,0,1] cc1_struct_0
s: 127
i: [1,130,0,6,3,0,0,255] cc4_struct_0
s: 130
i: [1,131,0,6,3,0,0,255] cc6_struct_0
s: 131
i: [1,147,0,6,2,0,0,255] t1_subset
s: 147
i: [1,152,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 152
i: [1,128,0,7,3,0,0,255] cc8_struct_0
s: 128
i: [1,132,0,7,3,0,0,255] cc11_struct_0
s: 132
i: [1,133,0,7,3,0,0,255] cc11_struct_0
s: 133
i: [1,134,0,7,3,0,0,255] cc9_struct_0
s: 134
i: [1,137,0,7,3,0,0,255] fc9_struct_0
s: 137
i: [1,141,0,7,3,0,0,1] fc2_struct_0
s: 141
i: [1,143,0,7,3,0,0,255] fc6_struct_0
s: 143
i: [1,144,0,7,3,0,0,255] fc8_struct_0
s: 144
i: [1,145,0,7,3,0,0,1] fc1_struct_0
s: 145
i: [1,146,0,7,3,0,0,255] fc7_struct_0
s: 146
i: [1,149,0,7,3,0,0,255] t8_boole
s: 149
d: [2,244,1,7,3,0,44,136,136]
s: 244
i: [1,148,0,8,3,0,0,255] t2_subset
s: 148
i: [1,135,0,9,4,0,0,255] cc10_struct_0
s: 135
d: [2,277,1,7,2,1,40,148,185]
a: [3,289,1,3,1,277]
s: 289
r: 289
i: [1,139,0,32,8,0,0,1] d4_diraf
s: 139
a: [3,288,1,3,1,277]
s: 288
i: [1,140,0,32,8,0,0,1] d4_diraf
s: 140
i: [1,142,0,34,9,0,0,1] t7_afvect0
s: 142
d: [2,310,2,4,2,1,40,289,141]
d: [2,311,2,2,1,2,21,310,237]
a: [3,314,2,2,1,311]
s: 314
i: [1,138,0,38,9,0,0,1] d4_diraf
s: 138
r: 288
r: 314
a: [3,315,2,2,1,311]
s: 315
d: [2,337,3,0,0,3,52,237,200,315,125]
e: 337
d: [2,338,3,4,2,1,40,315,125]
d: [2,341,3,2,1,2,21,338,200]
d: [2,342,3,0,0,3,21,341,237]
e: 342
r: 315
s: 288
a: [3,285,1,4,1,277]
s: 285
s: 314
d: [2,278,1,7,2,1,40,148,180]
a: [3,294,1,4,1,278]
s: 294
d: [2,279,1,7,2,1,40,148,175]
a: [3,299,1,4,1,279]
s: 299
d: [2,280,1,7,2,1,40,148,170]
a: [3,304,1,4,1,280]
s: 304
d: [2,320,1,32,8,1,40,138,165]
d: [2,321,1,28,7,2,21,320,170]
d: [2,322,1,24,6,3,21,321,175]
d: [2,323,1,20,5,4,21,322,180]
d: [2,324,1,16,4,5,21,323,185]
d: [2,325,1,14,3,6,21,324,190]
d: [2,326,1,12,2,7,21,325,314]
a: [3,330,1,6,1,326]
s: 330
r: 330
d: [2,348,2,4,1,1,40,285,152]
a: [3,354,2,4,1,348]
s: 354
a: [3,329,1,6,1,326]
s: 329
d: [2,356,2,4,1,1,40,294,152]
a: [3,362,2,4,1,356]
s: 362
a: [3,334,1,6,1,326]
s: 334
r: 329
r: 334
s: 330
d: [2,364,2,4,1,1,40,299,152]
a: [3,370,2,4,1,364]
s: 370
a: [3,333,1,6,1,326]
s: 333
d: [2,372,2,4,1,1,40,304,152]
a: [3,378,2,4,1,372]
s: 378
d: [2,281,1,6,2,0,40,148,151]
s: 281
d: [2,267,1,9,4,0,44,128,136]
s: 267
d: [2,380,2,28,8,1,40,330,142]
d: [2,383,2,24,7,2,21,380,175]
d: [2,384,2,20,6,3,21,383,170]
d: [2,385,2,16,5,4,21,384,180]
d: [2,386,2,12,4,5,21,385,185]
d: [2,387,2,10,3,6,21,386,190]
d: [2,388,2,8,2,7,21,387,195]
d: [2,389,2,6,1,8,21,388,200]
a: [3,393,2,6,1,389]
s: 393
d: [2,455,3,0,0,8,52,314,190,180,185,175,170,160,393,139]
e: 455
d: [2,458,3,26,7,1,40,393,139]
d: [2,466,3,20,6,2,21,458,160]
d: [2,467,3,16,5,3,21,466,170]
d: [2,468,3,12,4,4,21,467,175]
d: [2,469,3,8,3,5,21,468,185]
d: [2,470,3,4,2,6,21,469,180]
d: [2,471,3,2,1,7,21,470,190]
d: [2,472,3,0,0,8,21,471,314]
e: 472
r: 330
r: 333
r: 393
s: 329
s: 334
d: [2,269,1,9,4,0,44,134,136]
s: 269
a: [3,392,2,6,1,389]
s: 392
d: [2,382,2,26,7,1,40,330,139]
d: [2,395,2,22,6,2,21,382,175]
d: [2,396,2,18,5,3,21,395,170]
d: [2,397,2,14,4,4,21,396,180]
d: [2,398,2,10,3,5,21,397,185]
d: [2,399,2,8,2,6,21,398,190]
d: [2,400,2,6,1,7,21,399,314]
a: [3,404,2,6,1,400]
s: 404
d: [2,406,2,28,8,1,40,334,142]
d: [2,409,2,24,7,2,21,406,170]
d: [2,410,2,20,6,3,21,409,175]
d: [2,411,2,16,5,4,21,410,180]
d: [2,412,2,12,4,5,21,411,185]
d: [2,413,2,10,3,6,21,412,190]
d: [2,414,2,8,2,7,21,413,195]
d: [2,415,2,6,1,8,21,414,200]
a: [3,419,2,6,1,415]
s: 419
d: [2,499,3,0,0,8,52,314,190,180,185,175,170,160,419,140]
e: 499
d: [2,501,3,26,7,1,40,419,140]
d: [2,504,3,20,6,2,21,501,160]
d: [2,505,3,16,5,3,21,504,170]
d: [2,506,3,12,4,4,21,505,175]
d: [2,507,3,8,3,5,21,506,185]
d: [2,508,3,4,2,6,21,507,180]
d: [2,509,3,2,1,7,21,508,190]
d: [2,510,3,0,0,8,21,509,314]
e: 510
d: [2,428,2,28,8,1,40,330,142]
d: [2,431,2,24,7,2,21,428,175]
d: [2,432,2,20,6,3,21,431,170]
d: [2,433,2,16,5,4,21,432,180]
d: [2,434,2,12,4,5,21,433,185]
d: [2,435,2,10,3,6,21,434,190]
d: [2,436,2,8,2,7,21,435,195]
d: [2,437,2,6,1,8,21,436,200]
d: [2,474,2,28,8,1,40,334,142]
d: [2,477,2,24,7,2,21,474,170]
d: [2,478,2,20,6,3,21,477,175]
d: [2,479,2,16,5,4,21,478,180]
d: [2,480,2,12,4,5,21,479,185]
d: [2,481,2,10,3,6,21,480,190]
d: [2,482,2,8,2,7,21,481,195]
d: [2,483,2,6,1,8,21,482,200]
f: 510,483,472,437,342,325,232,110,111,112,113,114,115,116,117,118
% Refutation found. Thanks to Tanya!
% SZS status Theorem for afvect01__l3_afvect01
% SZS output start Proof for afvect01__l3_afvect01
1. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X3,X4) => r2_diraf(X0,X2,X1,X3,X4))))))) [input l3_afvect01]
2. ~! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X3,X4) => r2_diraf(X0,X2,X1,X3,X4))))))) [negated conjecture 1]
6. ! [X0] : (l1_struct_0(X0) => (~v7_struct_0(X0) => ~v2_struct_0(X0))) [input cc17_struct_0]
15. ! [X0] : ((l1_analoaf(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4)))))))) [input d4_diraf]
17. ! [X0] : (l1_analoaf(X0) => l1_struct_0(X0)) [input dt_l1_analoaf]
37. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (r2_analoaf(X0,X1,X2,X3,X4) => r2_analoaf(X0,X2,X1,X4,X3))))))) [input t7_afvect0]
44. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : ((~r2_diraf(X0,X2,X1,X3,X4) & r2_diraf(X0,X1,X2,X3,X4)) & m1_subset_1(X4,u1_struct_0(X0))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
45. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r2_diraf(X0,X2,X1,X3,X4) & r2_diraf(X0,X1,X2,X3,X4) & m1_subset_1(X4,u1_struct_0(X0))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) [flattening 44]
46. ! [X0] : (l1_struct_0(X0) | ~l1_analoaf(X0)) [ennf transformation 17]
51. ! [X0] : ((~v2_struct_0(X0) | v7_struct_0(X0)) | ~l1_struct_0(X0)) [ennf transformation 6]
52. ! [X0] : (~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0)) [flattening 51]
72. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | v2_struct_0(X0))) [ennf transformation 15]
73. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 72]
76. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_analoaf(X0,X2,X1,X4,X3) | ~r2_analoaf(X0,X1,X2,X3,X4)) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 37]
77. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (r2_analoaf(X0,X2,X1,X4,X3) | ~r2_analoaf(X0,X1,X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 76]
92. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r2_diraf(X0,X2,X1,X3,X4) & r2_diraf(X0,X1,X2,X3,X4) & m1_subset_1(X4,u1_struct_0(X0))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r2_diraf(sK0,X2,X1,X3,X4) & r2_diraf(sK0,X1,X2,X3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0)) [choice axiom]
93. ? [X1] : (? [X2] : (? [X3] : (? [X4] : (~r2_diraf(sK0,X2,X1,X3,X4) & r2_diraf(sK0,X1,X2,X3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (? [X4] : (~r2_diraf(sK0,X2,sK1,X3,X4) & r2_diraf(sK0,sK1,X2,X3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
94. ? [X2] : (? [X3] : (? [X4] : (~r2_diraf(sK0,X2,sK1,X3,X4) & r2_diraf(sK0,sK1,X2,X3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (? [X4] : (~r2_diraf(sK0,sK2,sK1,X3,X4) & r2_diraf(sK0,sK1,sK2,X3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
95. ? [X3] : (? [X4] : (~r2_diraf(sK0,sK2,sK1,X3,X4) & r2_diraf(sK0,sK1,sK2,X3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) & m1_subset_1(X3,u1_struct_0(sK0))) => (? [X4] : (~r2_diraf(sK0,sK2,sK1,sK3,X4) & r2_diraf(sK0,sK1,sK2,sK3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
96. ? [X4] : (~r2_diraf(sK0,sK2,sK1,sK3,X4) & r2_diraf(sK0,sK1,sK2,sK3,X4) & m1_subset_1(X4,u1_struct_0(sK0))) => (~r2_diraf(sK0,sK2,sK1,sK3,sK4) & r2_diraf(sK0,sK1,sK2,sK3,sK4) & m1_subset_1(sK4,u1_struct_0(sK0))) [choice axiom]
97. ((((~r2_diraf(sK0,sK2,sK1,sK3,sK4) & r2_diraf(sK0,sK1,sK2,sK3,sK4) & m1_subset_1(sK4,u1_struct_0(sK0))) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0) [skolemisation 45,96,95,94,93,92]
98. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r2_diraf(X0,X1,X2,X3,X4) | (~r2_analoaf(X0,X1,X2,X4,X3) & ~r2_analoaf(X0,X1,X2,X3,X4))) & ((r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4)) | ~r2_diraf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [nnf transformation 73]
99. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r2_diraf(X0,X1,X2,X3,X4) | (~r2_analoaf(X0,X1,X2,X4,X3) & ~r2_analoaf(X0,X1,X2,X3,X4))) & (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4) | ~r2_diraf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 98]
110. ~v7_struct_0(sK0) [cnf transformation 97] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
111. v1_afvect0(sK0) [cnf transformation 97] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
112. l1_analoaf(sK0) [cnf transformation 97] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
113. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 97] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
114. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 97] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
115. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 97] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
116. m1_subset_1(sK4,u1_struct_0(sK0)) [cnf transformation 97] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
117. r2_diraf(sK0,sK1,sK2,sK3,sK4) [cnf transformation 97] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
118. ~r2_diraf(sK0,sK2,sK1,sK3,sK4) [cnf transformation 97] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
121. ~l1_analoaf(X0) | l1_struct_0(X0) [cnf transformation 46] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
125. ~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 52] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
138. ~r2_diraf(X0,X1,X2,X3,X4) | r2_analoaf(X0,X1,X2,X3,X4) | r2_analoaf(X0,X1,X2,X4,X3) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 99] {a:0,w:38,wCS:15200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
139. ~r2_analoaf(X0,X1,X2,X3,X4) | r2_diraf(X0,X1,X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 99] {a:0,w:32,wCS:12800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
140. ~r2_analoaf(X0,X1,X2,X4,X3) | r2_diraf(X0,X1,X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 99] {a:0,w:32,wCS:12800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
142. ~r2_analoaf(X0,X1,X2,X3,X4) | r2_analoaf(X0,X2,X1,X4,X3) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 77] {a:0,w:34,wCS:13600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
158. 1 <=> r2_diraf(sK0,sK2,sK1,sK3,sK4) [avatar definition]
160. ~r2_diraf(sK0,sK2,sK1,sK3,sK4) <- (~1) [avatar component clause 158] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
161. ~1 [avatar split clause 118,158]
163. 2 <=> r2_diraf(sK0,sK1,sK2,sK3,sK4) [avatar definition]
165. r2_diraf(sK0,sK1,sK2,sK3,sK4) <- (2) [avatar component clause 163] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
166. 2 [avatar split clause 117,163]
168. 3 <=> m1_subset_1(sK4,u1_struct_0(sK0)) [avatar definition]
170. m1_subset_1(sK4,u1_struct_0(sK0)) <- (3) [avatar component clause 168] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
171. 3 [avatar split clause 116,168]
173. 4 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
175. m1_subset_1(sK3,u1_struct_0(sK0)) <- (4) [avatar component clause 173] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
176. 4 [avatar split clause 115,173]
178. 5 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
180. m1_subset_1(sK2,u1_struct_0(sK0)) <- (5) [avatar component clause 178] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
181. 5 [avatar split clause 114,178]
183. 6 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
185. m1_subset_1(sK1,u1_struct_0(sK0)) <- (6) [avatar component clause 183] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
186. 6 [avatar split clause 113,183]
188. 7 <=> l1_analoaf(sK0) [avatar definition]
190. l1_analoaf(sK0) <- (7) [avatar component clause 188] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
191. 7 [avatar split clause 112,188]
193. 8 <=> v1_afvect0(sK0) [avatar definition]
195. v1_afvect0(sK0) <- (8) [avatar component clause 193] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
196. 8 [avatar split clause 111,193]
198. 9 <=> v7_struct_0(sK0) [avatar definition]
200. ~v7_struct_0(sK0) <- (~9) [avatar component clause 198] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
201. ~9 [avatar split clause 110,198]
232. l1_struct_0(sK0) <- (7) [resolution 121,190] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
235. 16 <=> l1_struct_0(sK0) [avatar definition]
237. l1_struct_0(sK0) <- (16) [avatar component clause 235] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
238. 16 | ~7 [avatar split clause 232,188,235]
313. 27 <=> v2_struct_0(sK0) [avatar definition]
314. ~v2_struct_0(sK0) <- (~27) [avatar component clause 313] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
315. v2_struct_0(sK0) <- (27) [avatar component clause 313] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
320. r2_analoaf(sK0,sK1,sK2,sK3,sK4) | r2_analoaf(sK0,sK1,sK2,sK4,sK3) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (2) [resolution 138,165] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
321. r2_analoaf(sK0,sK1,sK2,sK3,sK4) | r2_analoaf(sK0,sK1,sK2,sK4,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (2, 3) [subsumption resolution 320,170] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
322. r2_analoaf(sK0,sK1,sK2,sK3,sK4) | r2_analoaf(sK0,sK1,sK2,sK4,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (2, 3, 4) [subsumption resolution 321,175] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
323. r2_analoaf(sK0,sK1,sK2,sK3,sK4) | r2_analoaf(sK0,sK1,sK2,sK4,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (2, 3, 4, 5) [subsumption resolution 322,180] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
324. r2_analoaf(sK0,sK1,sK2,sK3,sK4) | r2_analoaf(sK0,sK1,sK2,sK4,sK3) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (2, 3, 4, 5, 6) [subsumption resolution 323,185] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
325. r2_analoaf(sK0,sK1,sK2,sK3,sK4) | r2_analoaf(sK0,sK1,sK2,sK4,sK3) | v2_struct_0(sK0) <- (2, 3, 4, 5, 6, 7) [subsumption resolution 324,190] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
328. 28 <=> r2_analoaf(sK0,sK1,sK2,sK4,sK3) [avatar definition]
330. r2_analoaf(sK0,sK1,sK2,sK4,sK3) <- (28) [avatar component clause 328] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
332. 29 <=> r2_analoaf(sK0,sK1,sK2,sK3,sK4) [avatar definition]
334. r2_analoaf(sK0,sK1,sK2,sK3,sK4) <- (29) [avatar component clause 332] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
336. 27 | 28 | 29 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 325,188,183,178,173,168,163,332,328,313]
338. v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (27) [resolution 315,125] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
341. ~l1_struct_0(sK0) <- (~9, 27) [subsumption resolution 338,200] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
342. $false <- (~9, 16, 27) [subsumption resolution 341,237] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
343. 9 | ~16 | ~27 [avatar contradiction clause 342]
391. 34 <=> r2_analoaf(sK0,sK2,sK1,sK3,sK4) [avatar definition]
393. r2_analoaf(sK0,sK2,sK1,sK3,sK4) <- (34) [avatar component clause 391] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
417. 36 <=> r2_analoaf(sK0,sK2,sK1,sK4,sK3) [avatar definition]
419. r2_analoaf(sK0,sK2,sK1,sK4,sK3) <- (36) [avatar component clause 417] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
428. r2_analoaf(sK0,sK2,sK1,sK3,sK4) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (28) [resolution 330,142] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
431. r2_analoaf(sK0,sK2,sK1,sK3,sK4) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (4, 28) [subsumption resolution 428,175] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
432. r2_analoaf(sK0,sK2,sK1,sK3,sK4) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (3, 4, 28) [subsumption resolution 431,170] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
433. r2_analoaf(sK0,sK2,sK1,sK3,sK4) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 28) [subsumption resolution 432,180] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
434. r2_analoaf(sK0,sK2,sK1,sK3,sK4) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 6, 28) [subsumption resolution 433,185] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
435. r2_analoaf(sK0,sK2,sK1,sK3,sK4) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 6, 7, 28) [subsumption resolution 434,190] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
436. r2_analoaf(sK0,sK2,sK1,sK3,sK4) | v7_struct_0(sK0) <- (3, 4, 5, 6, 7, 8, 28) [subsumption resolution 435,195] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
437. r2_analoaf(sK0,sK2,sK1,sK3,sK4) <- (3, 4, 5, 6, 7, 8, ~9, 28) [subsumption resolution 436,200] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
438. 34 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | 9 | ~28 [avatar split clause 437,328,198,193,188,183,178,173,168,391]
458. r2_diraf(sK0,sK2,sK1,sK3,sK4) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (34) [resolution 393,139] {a:3,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
466. ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 34) [subsumption resolution 458,160] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
467. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 34) [subsumption resolution 466,170] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
468. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 34) [subsumption resolution 467,175] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
469. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 6, 34) [subsumption resolution 468,185] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
470. ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 5, 6, 34) [subsumption resolution 469,180] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
471. v2_struct_0(sK0) <- (~1, 3, 4, 5, 6, 7, 34) [subsumption resolution 470,190] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
472. $false <- (~1, 3, 4, 5, 6, 7, ~27, 34) [subsumption resolution 471,314] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
473. 1 | ~3 | ~4 | ~5 | ~6 | ~7 | 27 | ~34 [avatar contradiction clause 472]
474. r2_analoaf(sK0,sK2,sK1,sK4,sK3) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (29) [resolution 334,142] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
477. r2_analoaf(sK0,sK2,sK1,sK4,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (3, 29) [subsumption resolution 474,170] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
478. r2_analoaf(sK0,sK2,sK1,sK4,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (3, 4, 29) [subsumption resolution 477,175] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
479. r2_analoaf(sK0,sK2,sK1,sK4,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 29) [subsumption resolution 478,180] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
480. r2_analoaf(sK0,sK2,sK1,sK4,sK3) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 6, 29) [subsumption resolution 479,185] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
481. r2_analoaf(sK0,sK2,sK1,sK4,sK3) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (3, 4, 5, 6, 7, 29) [subsumption resolution 480,190] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
482. r2_analoaf(sK0,sK2,sK1,sK4,sK3) | v7_struct_0(sK0) <- (3, 4, 5, 6, 7, 8, 29) [subsumption resolution 481,195] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
483. r2_analoaf(sK0,sK2,sK1,sK4,sK3) <- (3, 4, 5, 6, 7, 8, ~9, 29) [subsumption resolution 482,200] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
484. 36 | ~3 | ~4 | ~5 | ~6 | ~7 | ~8 | 9 | ~29 [avatar split clause 483,332,198,193,188,183,178,173,168,417]
501. r2_diraf(sK0,sK2,sK1,sK3,sK4) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (36) [resolution 419,140] {a:3,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
504. ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 36) [subsumption resolution 501,160] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
505. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 36) [subsumption resolution 504,170] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
506. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 36) [subsumption resolution 505,175] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
507. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 6, 36) [subsumption resolution 506,185] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
508. ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 5, 6, 36) [subsumption resolution 507,180] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
509. v2_struct_0(sK0) <- (~1, 3, 4, 5, 6, 7, 36) [subsumption resolution 508,190] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
510. $false <- (~1, 3, 4, 5, 6, 7, ~27, 36) [subsumption resolution 509,314] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
511. 1 | ~3 | ~4 | ~5 | ~6 | ~7 | 27 | ~36 [avatar contradiction clause 510]
518. $false [avatar sat refutation 161,166,171,176,181,186,191,196,201,238,336,343,438,473,484,511] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:538,thDist:-538}
% SZS output end Proof for afvect01__l3_afvect01
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% Introduced skolems: 10
% 
% >>> Saturation
% Initial clauses: 47
% Generated clauses: 264
% Active clauses: 73
% Passive clauses: 80
% Final active clauses: 63
% Final passive clauses: 1
% Main loop iterations started: 78
% 
% >>> Simplifying Inferences
% Duplicate literals: 15
% Fw subsumption resolutions: 90
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 25
% 
% >>> Generating Inferences
% Binary resolution: 62
% Unit resulting resolution: 3
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 17
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 48
% SAT solver unit clauses: 15
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

