i: [1,181,0,2,1,1,0,0] l18_orders_3
a: [3,282,0,2,1,181]
s: 282
i: [1,180,0,2,1,1,0,0] l18_orders_3
a: [3,287,0,2,1,180]
s: 287
i: [1,182,0,3,1,1,0,0] l18_orders_3
a: [3,277,0,3,1,182]
s: 277
i: [1,183,0,5,1,1,0,0] l18_orders_3
a: [3,272,0,5,1,183]
s: 272
i: [1,184,0,2,1,0,0,255] spc1_boole
a: [3,292,0,2,1,184]
s: 292
i: [1,185,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,297,0,2,1,185]
s: 297
i: [1,186,0,2,1,0,0,3] fc1_xboole_0
a: [3,302,0,2,1,186]
s: 302
i: [1,249,0,2,1,0,0,1] rc2_xboole_0
a: [3,312,0,2,1,249]
s: 312
i: [1,250,0,2,1,0,0,2] existence_l1_struct_0
a: [3,317,0,2,1,250]
s: 317
i: [1,251,0,2,1,0,0,255] existence_l1_orders_2
a: [3,322,0,2,1,251]
s: 322
i: [1,252,0,2,1,0,0,1] rc1_xboole_0
a: [3,327,0,2,1,252]
s: 327
i: [1,254,0,2,1,0,0,255] rc7_funct_1
a: [3,332,0,2,1,254]
s: 332
i: [1,253,0,2,1,0,0,255] rc7_funct_1
a: [3,337,0,2,1,253]
s: 337
i: [1,256,0,2,1,0,0,1] rc4_orders_3
a: [3,342,0,2,1,256]
s: 342
i: [1,255,0,2,1,0,0,1] rc4_orders_3
a: [3,347,0,2,1,255]
s: 347
i: [1,258,0,2,1,0,0,255] rc1_funct_1
a: [3,352,0,2,1,258]
s: 352
i: [1,257,0,2,1,0,0,255] rc1_funct_1
a: [3,357,0,2,1,257]
s: 357
i: [1,261,0,2,1,0,0,255] rc5_funct_1
a: [3,362,0,2,1,261]
s: 362
i: [1,260,0,2,1,0,0,255] rc5_funct_1
a: [3,367,0,2,1,260]
s: 367
i: [1,259,0,2,1,0,0,255] rc5_funct_1
a: [3,372,0,2,1,259]
s: 372
i: [1,263,0,2,1,0,0,255] rc2_funct_1
a: [3,377,0,2,1,263]
s: 377
i: [1,262,0,2,1,0,0,255] rc2_funct_1
a: [3,382,0,2,1,262]
s: 382
i: [1,187,0,3,1,0,0,255] d2_xboole_0
a: [3,307,0,3,1,187]
s: 307
r: 302
i: [1,204,0,4,2,0,0,255] dt_l1_orders_2
s: 204
i: [1,205,0,4,2,0,0,255] cc7_funct_1
s: 205
i: [1,206,0,4,2,0,0,255] cc1_funct_1
s: 206
d: [2,385,1,2,1,1,40,204,322]
a: [3,389,1,2,1,385]
s: 389
i: [1,247,0,4,1,0,0,2] existence_m1_subset_1
s: 247
d: [2,391,1,2,1,1,40,205,297]
a: [3,396,1,2,1,391]
s: 396
i: [1,207,0,5,2,0,0,2] t6_boole
s: 207
i: [1,246,0,5,2,0,0,1] t7_boole
s: 246
d: [2,392,1,2,1,1,40,205,327]
a: [3,401,1,2,1,392]
s: 401
d: [2,415,0,5,2,1,22,207,307]
s: 415
d: [2,403,1,2,1,1,40,206,297]
a: [3,408,1,2,1,403]
s: 408
i: [1,190,0,6,3,0,0,255] cc7_struct_0
s: 190
i: [1,191,0,6,3,0,0,255] cc5_struct_0
s: 191
d: [2,404,1,2,1,1,40,206,327]
a: [3,413,1,2,1,404]
s: 413
i: [1,193,0,6,3,0,0,255] cc17_struct_0
s: 193
i: [1,194,0,6,3,0,0,255] cc2_struct_0
s: 194
i: [1,195,0,6,3,0,0,255] cc1_struct_0
s: 195
i: [1,198,0,6,3,0,0,255] cc4_struct_0
s: 198
i: [1,199,0,6,3,0,0,255] cc6_struct_0
s: 199
i: [1,235,0,6,2,0,0,1] t1_subset
s: 235
i: [1,248,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 248
i: [1,188,0,7,3,0,0,255] cc8_funct_1
s: 188
i: [1,189,0,7,3,0,0,255] cc8_funct_1
s: 189
d: [2,420,1,5,2,0,40,188,247]
s: 420
i: [1,196,0,7,3,0,0,255] cc8_struct_0
s: 196
d: [2,421,1,5,2,0,40,189,247]
s: 421
i: [1,200,0,7,3,0,0,255] cc11_struct_0
s: 200
i: [1,201,0,7,3,0,0,255] cc11_struct_0
s: 201
i: [1,202,0,7,3,0,0,255] cc9_struct_0
s: 202
i: [1,208,0,7,3,0,0,1] fc2_struct_0
s: 208
i: [1,209,0,7,3,0,0,255] fc9_struct_0
s: 209
i: [1,210,0,7,3,0,0,255] fc6_struct_0
s: 210
i: [1,224,0,7,3,0,0,1] fc1_struct_0
s: 224
i: [1,225,0,7,3,0,0,255] fc8_struct_0
s: 225
i: [1,226,0,7,3,0,0,255] fc7_struct_0
s: 226
i: [1,245,0,7,3,0,0,2] t8_boole
s: 245
d: [2,422,0,7,3,1,22,196,307]
s: 422
d: [2,423,0,7,3,1,22,202,307]
s: 423
d: [2,416,1,7,3,1,44,415,415]
s: 416
i: [1,219,0,8,3,0,0,1] existence_m1_orders_3
s: 219
d: [2,425,1,7,3,0,40,224,206]
s: 425
i: [1,220,0,8,3,0,0,1] existence_m1_orders_3
s: 220
i: [1,223,0,8,4,0,0,255] cc6_funct_1
s: 223
d: [2,426,1,7,3,0,40,224,205]
s: 426
i: [1,227,0,8,4,0,0,255] cc5_funct_1
s: 227
i: [1,234,0,8,4,0,0,255] cc4_funct_1
s: 234
d: [2,437,0,6,3,0,21,234,206]
s: 437
i: [1,236,0,8,3,0,0,1] t2_subset
s: 236
d: [2,439,1,4,2,1,40,437,357]
a: [3,445,1,2,1,439]
s: 445
i: [1,268,0,8,2,0,0,1] d7_orders_3
s: 268
d: [2,440,1,4,2,1,40,437,372]
d: [2,451,1,2,1,2,21,440,362]
a: [3,455,1,2,1,451]
s: 455
i: [1,203,0,9,4,0,0,255] cc10_struct_0
s: 203
i: [1,211,0,9,4,0,0,255] dt_m1_orders_3
s: 211
d: [2,471,1,6,3,1,40,211,277]
d: [2,476,1,4,2,2,21,471,282]
d: [2,477,1,2,1,3,21,476,287]
a: [3,481,1,2,1,477]
s: 481
i: [1,212,0,9,4,0,0,255] dt_m1_orders_3
s: 212
d: [2,483,1,6,3,1,40,212,277]
d: [2,488,1,4,2,2,21,483,282]
d: [2,489,1,2,1,3,21,488,287]
a: [3,493,1,2,1,489]
s: 493
i: [1,213,0,9,4,0,0,255] dt_m1_orders_3
s: 213
i: [1,214,0,9,4,0,0,255] dt_m1_orders_3
s: 214
d: [2,495,2,2,1,1,40,493,204]
a: [3,499,2,2,1,495]
s: 499
i: [1,267,0,9,2,0,0,1] d7_orders_3
s: 267
d: [2,441,1,4,2,1,40,437,382]
a: [3,460,1,2,1,441]
s: 460
i: [1,215,0,10,4,0,0,2] redefinition_m1_orders_3
s: 215
i: [1,216,0,10,4,0,0,2] redefinition_m1_orders_3
s: 216
d: [2,505,1,7,3,1,40,215,277]
d: [2,510,1,5,2,2,21,505,282]
d: [2,511,1,3,1,3,21,510,287]
a: [3,515,1,3,1,511]
s: 515
i: [1,217,0,10,4,0,0,2] redefinition_m1_orders_3
s: 217
d: [2,520,2,4,2,1,40,515,189]
a: [3,531,2,2,1,520]
s: 531
i: [1,218,0,10,4,0,0,2] redefinition_m1_orders_3
s: 218
i: [1,265,0,10,3,0,0,1] d7_orders_3
s: 265
d: [2,519,2,5,2,1,40,515,236]
d: [2,522,2,3,1,2,21,519,287]
d: [2,542,1,5,2,1,40,265,272]
d: [2,547,1,3,1,2,21,542,499]
f: 547,522,511,495,489,180,181,182,183
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_3__l18_orders_3
% SZS output start Proof for orders_3__l18_orders_3
1. ! [X0] : ((v4_orders_3(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_orders_3(X1,X0) => r2_hidden(u1_struct_0(X1),k2_orders_3(X0)))) [input l18_orders_3]
2. ~! [X0] : ((v4_orders_3(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_orders_3(X1,X0) => r2_hidden(u1_struct_0(X1),k2_orders_3(X0)))) [negated conjecture 1]
27. ! [X0] : (l1_orders_2(X0) => l1_struct_0(X0)) [input dt_l1_orders_2]
39. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
45. ! [X0] : ((v4_orders_3(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_orders_3(X1,X0) <=> m1_subset_1(X1,X0))) [input redefinition_m1_orders_3]
51. ! [X0] : ((v4_orders_3(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_orders_3(X1,X0) => (l1_orders_2(X1) & v5_orders_2(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~v2_struct_0(X1)))) [input dt_m1_orders_3]
62. ! [X0] : ! [X1] : (k2_orders_3(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (u1_struct_0(X3) = X2 & r2_hidden(X3,X0) & l1_struct_0(X3)))) [input d7_orders_3]
64. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 39]
65. ! [X0,X1] : (k2_orders_3(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (u1_struct_0(X3) = X2 & r2_hidden(X3,X0) & l1_struct_0(X3)))) [flattening 62]
69. ! [X0] : ((v4_orders_3(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_orders_3(X1,X0) => (l1_orders_2(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~v2_struct_0(X1)))) [pure predicate removal 51]
71. ! [X0] : ((v4_orders_3(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_orders_3(X1,X0) => (l1_orders_2(X1) & v3_orders_2(X1) & ~v2_struct_0(X1)))) [pure predicate removal 69]
75. ! [X0] : ((v4_orders_3(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_orders_3(X1,X0) => (l1_orders_2(X1) & ~v2_struct_0(X1)))) [pure predicate removal 71]
78. ? [X0] : (? [X1] : (~r2_hidden(u1_struct_0(X1),k2_orders_3(X0)) & m1_orders_3(X1,X0)) & (v4_orders_3(X0) & ~v1_xboole_0(X0))) [ennf transformation 2]
79. ? [X0] : (? [X1] : (~r2_hidden(u1_struct_0(X1),k2_orders_3(X0)) & m1_orders_3(X1,X0)) & v4_orders_3(X0) & ~v1_xboole_0(X0)) [flattening 78]
103. ! [X0] : (l1_struct_0(X0) | ~l1_orders_2(X0)) [ennf transformation 27]
113. ! [X0] : (! [X1] : ((l1_orders_2(X1) & ~v2_struct_0(X1)) | ~m1_orders_3(X1,X0)) | (~v4_orders_3(X0) | v1_xboole_0(X0))) [ennf transformation 75]
114. ! [X0] : (! [X1] : ((l1_orders_2(X1) & ~v2_struct_0(X1)) | ~m1_orders_3(X1,X0)) | ~v4_orders_3(X0) | v1_xboole_0(X0)) [flattening 113]
117. ! [X0] : (! [X1] : (m1_orders_3(X1,X0) <=> m1_subset_1(X1,X0)) | (~v4_orders_3(X0) | v1_xboole_0(X0))) [ennf transformation 45]
118. ! [X0] : (! [X1] : (m1_orders_3(X1,X0) <=> m1_subset_1(X1,X0)) | ~v4_orders_3(X0) | v1_xboole_0(X0)) [flattening 117]
140. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 64]
141. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 140]
145. ? [X0] : (? [X1] : (~r2_hidden(u1_struct_0(X1),k2_orders_3(X0)) & m1_orders_3(X1,X0)) & v4_orders_3(X0) & ~v1_xboole_0(X0)) => (? [X1] : (~r2_hidden(u1_struct_0(X1),k2_orders_3(sK0)) & m1_orders_3(X1,sK0)) & v4_orders_3(sK0) & ~v1_xboole_0(sK0)) [choice axiom]
146. ? [X1] : (~r2_hidden(u1_struct_0(X1),k2_orders_3(sK0)) & m1_orders_3(X1,sK0)) => (~r2_hidden(u1_struct_0(sK1),k2_orders_3(sK0)) & m1_orders_3(sK1,sK0)) [choice axiom]
147. (~r2_hidden(u1_struct_0(sK1),k2_orders_3(sK0)) & m1_orders_3(sK1,sK0)) & v4_orders_3(sK0) & ~v1_xboole_0(sK0) [skolemisation 79,146,145]
148. ! [X0] : (! [X1] : ((m1_orders_3(X1,X0) | ~m1_subset_1(X1,X0)) & (m1_subset_1(X1,X0) | ~m1_orders_3(X1,X0))) | ~v4_orders_3(X0) | v1_xboole_0(X0)) [nnf transformation 118]
154. ! [X0,X1] : ((k2_orders_3(X0) = X1 | ? [X2] : ((! [X3] : (u1_struct_0(X3) != X2 | ~r2_hidden(X3,X0) | ~l1_struct_0(X3)) | ~r2_hidden(X2,X1)) & (? [X3] : (u1_struct_0(X3) = X2 & r2_hidden(X3,X0) & l1_struct_0(X3)) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ! [X3] : (u1_struct_0(X3) != X2 | ~r2_hidden(X3,X0) | ~l1_struct_0(X3))) & (? [X3] : (u1_struct_0(X3) = X2 & r2_hidden(X3,X0) & l1_struct_0(X3)) | ~r2_hidden(X2,X1))) | k2_orders_3(X0) != X1)) [nnf transformation 65]
155. ! [X0,X1] : ((k2_orders_3(X0) = X1 | ? [X2] : ((! [X3] : (u1_struct_0(X3) != X2 | ~r2_hidden(X3,X0) | ~l1_struct_0(X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (u1_struct_0(X4) = X2 & r2_hidden(X4,X0) & l1_struct_0(X4)) | r2_hidden(X2,X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (u1_struct_0(X6) != X5 | ~r2_hidden(X6,X0) | ~l1_struct_0(X6))) & (? [X7] : (u1_struct_0(X7) = X5 & r2_hidden(X7,X0) & l1_struct_0(X7)) | ~r2_hidden(X5,X1))) | k2_orders_3(X0) != X1)) [rectify 154]
156. ! [X1,X0] : (? [X2] : ((! [X3] : (u1_struct_0(X3) != X2 | ~r2_hidden(X3,X0) | ~l1_struct_0(X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (u1_struct_0(X4) = X2 & r2_hidden(X4,X0) & l1_struct_0(X4)) | r2_hidden(X2,X1))) => ((! [X3] : (u1_struct_0(X3) != sK4(X0,X1) | ~r2_hidden(X3,X0) | ~l1_struct_0(X3)) | ~r2_hidden(sK4(X0,X1),X1)) & (? [X4] : (u1_struct_0(X4) = sK4(X0,X1) & r2_hidden(X4,X0) & l1_struct_0(X4)) | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
157. ! [X1,X0] : (? [X4] : (u1_struct_0(X4) = sK4(X0,X1) & r2_hidden(X4,X0) & l1_struct_0(X4)) => (sK4(X0,X1) = u1_struct_0(sK5(X0,X1)) & r2_hidden(sK5(X0,X1),X0) & l1_struct_0(sK5(X0,X1)))) [choice axiom]
158. ! [X5,X0] : (? [X7] : (u1_struct_0(X7) = X5 & r2_hidden(X7,X0) & l1_struct_0(X7)) => (u1_struct_0(sK6(X0,X5)) = X5 & r2_hidden(sK6(X0,X5),X0) & l1_struct_0(sK6(X0,X5)))) [choice axiom]
159. ! [X0,X1] : ((k2_orders_3(X0) = X1 | ((! [X3] : (u1_struct_0(X3) != sK4(X0,X1) | ~r2_hidden(X3,X0) | ~l1_struct_0(X3)) | ~r2_hidden(sK4(X0,X1),X1)) & ((sK4(X0,X1) = u1_struct_0(sK5(X0,X1)) & r2_hidden(sK5(X0,X1),X0) & l1_struct_0(sK5(X0,X1))) | r2_hidden(sK4(X0,X1),X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (u1_struct_0(X6) != X5 | ~r2_hidden(X6,X0) | ~l1_struct_0(X6))) & ((u1_struct_0(sK6(X0,X5)) = X5 & r2_hidden(sK6(X0,X5),X0) & l1_struct_0(sK6(X0,X5))) | ~r2_hidden(X5,X1))) | k2_orders_3(X0) != X1)) [skolemisation 155,158,157,156]
180. ~v1_xboole_0(sK0) [cnf transformation 147] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
181. v4_orders_3(sK0) [cnf transformation 147] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
182. m1_orders_3(sK1,sK0) [cnf transformation 147] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
183. ~r2_hidden(u1_struct_0(sK1),k2_orders_3(sK0)) [cnf transformation 147] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
204. ~l1_orders_2(X0) | l1_struct_0(X0) [cnf transformation 103] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
212. ~m1_orders_3(X1,X0) | l1_orders_2(X1) | ~v4_orders_3(X0) | v1_xboole_0(X0) [cnf transformation 114] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
215. ~m1_orders_3(X1,X0) | m1_subset_1(X1,X0) | ~v4_orders_3(X0) | v1_xboole_0(X0) [cnf transformation 148] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
236. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 141] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
240. r2_hidden(X5,X1) | u1_struct_0(X6) != X5 | ~r2_hidden(X6,X0) | ~l1_struct_0(X6) | k2_orders_3(X0) != X1 [cnf transformation 159] {a:0,w:16,wCS:6400,sine:1,thAx:0,allAx:4,thDist:-4}
264. r2_hidden(u1_struct_0(X6),X1) | ~r2_hidden(X6,X0) | ~l1_struct_0(X6) | k2_orders_3(X0) != X1 [equality resolution 240] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:4,thDist:-4}
265. r2_hidden(u1_struct_0(X6),k2_orders_3(X0)) | ~r2_hidden(X6,X0) | ~l1_struct_0(X6) [equality resolution 264] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
270. 1 <=> r2_hidden(u1_struct_0(sK1),k2_orders_3(sK0)) [avatar definition]
272. ~r2_hidden(u1_struct_0(sK1),k2_orders_3(sK0)) <- (~1) [avatar component clause 270] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
273. ~1 [avatar split clause 183,270]
275. 2 <=> m1_orders_3(sK1,sK0) [avatar definition]
277. m1_orders_3(sK1,sK0) <- (2) [avatar component clause 275] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
278. 2 [avatar split clause 182,275]
280. 3 <=> v4_orders_3(sK0) [avatar definition]
282. v4_orders_3(sK0) <- (3) [avatar component clause 280] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
283. 3 [avatar split clause 181,280]
285. 4 <=> v1_xboole_0(sK0) [avatar definition]
287. ~v1_xboole_0(sK0) <- (~4) [avatar component clause 285] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
288. ~4 [avatar split clause 180,285]
483. l1_orders_2(sK1) | ~v4_orders_3(sK0) | v1_xboole_0(sK0) <- (2) [resolution 212,277] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
488. l1_orders_2(sK1) | v1_xboole_0(sK0) <- (2, 3) [subsumption resolution 483,282] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
489. l1_orders_2(sK1) <- (2, 3, ~4) [subsumption resolution 488,287] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
491. 35 <=> l1_orders_2(sK1) [avatar definition]
493. l1_orders_2(sK1) <- (35) [avatar component clause 491] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
494. 35 | ~2 | ~3 | 4 [avatar split clause 489,285,280,275,491]
495. l1_struct_0(sK1) <- (35) [resolution 493,204] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
497. 36 <=> l1_struct_0(sK1) [avatar definition]
499. l1_struct_0(sK1) <- (36) [avatar component clause 497] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
500. 36 | ~35 [avatar split clause 495,491,497]
505. m1_subset_1(sK1,sK0) | ~v4_orders_3(sK0) | v1_xboole_0(sK0) <- (2) [resolution 215,277] {a:1,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
510. m1_subset_1(sK1,sK0) | v1_xboole_0(sK0) <- (2, 3) [subsumption resolution 505,282] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
511. m1_subset_1(sK1,sK0) <- (2, 3, ~4) [subsumption resolution 510,287] {a:1,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
513. 37 <=> m1_subset_1(sK1,sK0) [avatar definition]
515. m1_subset_1(sK1,sK0) <- (37) [avatar component clause 513] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
516. 37 | ~2 | ~3 | 4 [avatar split clause 511,285,280,275,513]
519. v1_xboole_0(sK0) | r2_hidden(sK1,sK0) <- (37) [resolution 515,236] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
522. r2_hidden(sK1,sK0) <- (~4, 37) [subsumption resolution 519,287] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
524. 38 <=> r2_hidden(sK1,sK0) [avatar definition]
527. 38 | 4 | ~37 [avatar split clause 522,513,285,524]
542. ~r2_hidden(sK1,sK0) | ~l1_struct_0(sK1) <- (~1) [resolution 265,272] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
547. ~r2_hidden(sK1,sK0) <- (~1, 36) [subsumption resolution 542,499] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
548. ~38 | 1 | ~36 [avatar split clause 547,497,270,524]
549. $false [avatar sat refutation 273,278,283,288,494,500,516,527,548] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:175,thDist:-175}
% SZS output end Proof for orders_3__l18_orders_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 61
% >>> Preprocessing
% Introduced skolems: 17
% Pure predicates: 6
% 
% >>> Saturation
% Initial clauses: 84
% Generated clauses: 198
% Active clauses: 77
% Passive clauses: 117
% Final active clauses: 76
% Final passive clauses: 26
% Main loop iterations started: 91
% 
% >>> Simplifying Inferences
% Duplicate literals: 18
% Fw subsumption resolutions: 10
% Fw demodulations: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Forward subsumptions: 14
% 
% >>> Generating Inferences
% Binary resolution: 45
% Forward superposition: 3
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 41
% SAT solver unit clauses: 23
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

