i: [1,233,0,2,1,1,0,0] t44_bciideal
a: [3,2843,0,2,1,233]
s: 2843
i: [1,232,0,2,1,1,0,0] t44_bciideal
a: [3,2848,0,2,1,232]
s: 2848
i: [1,231,0,2,1,1,0,0] t44_bciideal
a: [3,2853,0,2,1,231]
s: 2853
i: [1,230,0,2,1,1,0,0] t44_bciideal
a: [3,2858,0,2,1,230]
s: 2858
i: [1,229,0,2,1,1,0,0] t44_bciideal
a: [3,2863,0,2,1,229]
s: 2863
i: [1,228,0,2,1,1,0,0] t44_bciideal
a: [3,2868,0,2,1,228]
s: 2868
i: [1,227,0,2,1,1,0,0] t44_bciideal
a: [3,2873,0,2,1,227]
s: 2873
i: [1,1433,0,85,6,1,0,0] t44_bciideal
a: [3,1575,0,4,1,1433]
s: 1575
i: [1,1432,0,85,6,1,0,0] t44_bciideal
a: [3,1580,0,4,1,1432]
s: 1580
i: [1,1355,0,68,6,1,0,0] t44_bciideal
a: [3,1689,0,4,1,1355]
s: 1689
i: [1,1275,0,68,6,1,0,0] t44_bciideal
a: [3,1773,0,4,1,1275]
s: 1773
i: [1,235,0,5,1,1,0,0] t44_bciideal
a: [3,2833,0,5,1,235]
s: 2833
i: [1,234,0,5,1,1,0,0] t44_bciideal
a: [3,2838,0,5,1,234]
s: 2838
i: [1,1435,0,97,7,1,0,0] t44_bciideal
a: [3,1545,0,6,1,1435]
s: 1545
i: [1,1436,0,2,1,0,0,255] spc1_boole
a: [3,2878,0,2,1,1436]
s: 2878
i: [1,1437,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,2883,0,2,1,1437]
s: 2883
i: [1,1518,0,2,1,0,0,255] existence_l1_bcialg_1
a: [3,2893,0,2,1,1518]
s: 2893
i: [1,1519,0,2,1,0,0,2] existence_l2_struct_0
a: [3,2898,0,2,1,1519]
s: 2898
i: [1,1520,0,2,1,0,0,1] existence_l2_bcialg_1
a: [3,2903,0,2,1,1520]
s: 2903
i: [1,1521,0,2,1,0,0,3] existence_l1_struct_0
a: [3,2908,0,2,1,1521]
s: 2908
i: [1,1524,0,2,1,0,0,2] rc20_struct_0
a: [3,2913,0,2,1,1524]
s: 2913
i: [1,1523,0,2,1,0,0,2] rc20_struct_0
a: [3,2918,0,2,1,1523]
s: 2918
i: [1,1522,0,2,1,0,0,2] rc20_struct_0
a: [3,2923,0,2,1,1522]
s: 2923
i: [1,1532,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2928,0,2,1,1532]
s: 2928
i: [1,1531,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2933,0,2,1,1531]
s: 2933
i: [1,1530,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2938,0,2,1,1530]
s: 2938
i: [1,1529,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2943,0,2,1,1529]
s: 2943
i: [1,1528,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2948,0,2,1,1528]
s: 2948
i: [1,1527,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2953,0,2,1,1527]
s: 2953
i: [1,1526,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2958,0,2,1,1526]
s: 2958
i: [1,1525,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2963,0,2,1,1525]
s: 2963
i: [1,1513,0,3,1,0,0,255] reflexivity_r1_tarski
s: 1513
i: [1,1438,0,3,1,0,0,255] d2_xboole_0
a: [3,2888,0,3,1,1438]
s: 2888
r: 2883
d: [2,2965,0,2,1,2,23,2883,2888]
a: [3,2969,0,2,1,2965]
s: 2969
a: [3,1561,0,15,1,1435]
s: 1561
i: [1,1439,0,4,2,0,0,255] dt_l1_bcialg_1
s: 1439
d: [2,2971,1,2,1,1,40,1439,2893]
a: [3,2975,1,2,1,2971]
s: 2975
i: [1,1444,0,4,2,0,0,2] dt_l2_struct_0
s: 1444
d: [2,2977,1,2,1,1,40,1444,2898]
a: [3,2982,1,2,1,2977]
s: 2982
i: [1,1450,0,4,2,0,0,255] dt_l2_bcialg_1
s: 1450
i: [1,1451,0,4,2,0,0,255] dt_l2_bcialg_1
s: 1451
d: [2,2989,1,2,1,1,40,1450,2843]
a: [3,2995,1,2,1,2989]
s: 2995
i: [1,1512,0,4,1,0,0,1] existence_m1_subset_1
s: 1512
d: [2,3007,1,2,1,1,40,1451,2843]
a: [3,3013,1,2,1,3007]
s: 3013
i: [1,1440,0,5,2,0,0,255] dt_u1_bcialg_1
s: 1440
i: [1,1449,0,5,2,0,0,255] t6_boole
s: 1449
d: [2,3025,2,2,1,1,40,2995,1439]
a: [3,3029,2,2,1,3025]
s: 3029
i: [1,1509,0,5,2,0,0,255] t7_boole
s: 1509
d: [2,2978,1,2,1,1,40,1444,2923]
a: [3,2987,1,2,1,2978]
s: 2987
i: [1,1445,0,6,2,0,0,255] fc10_struct_0
s: 1445
i: [1,1448,0,6,2,0,0,255] rc17_struct_0
s: 1448
d: [2,2990,1,2,1,1,40,1450,2903]
a: [3,3000,1,2,1,2990]
s: 3000
i: [1,1452,0,6,3,0,0,255] cc7_struct_0
s: 1452
d: [2,2991,1,2,1,1,40,1450,2963]
a: [3,3005,1,2,1,2991]
s: 3005
i: [1,1453,0,6,3,0,0,255] cc5_struct_0
s: 1453
i: [1,1455,0,6,3,0,0,3] cc2_struct_0
s: 1455
d: [2,3008,1,2,1,1,40,1451,2903]
a: [3,3018,1,2,1,3008]
s: 3018
i: [1,1456,0,6,3,0,0,3] cc17_struct_0
s: 1456
d: [2,3009,1,2,1,1,40,1451,2963]
a: [3,3023,1,2,1,3009]
s: 3023
i: [1,1457,0,6,3,0,0,255] cc6_struct_0
s: 1457
i: [1,1458,0,6,3,0,0,3] cc1_struct_0
s: 1458
d: [2,3041,1,4,2,1,40,1452,2913]
d: [2,3042,1,2,1,2,21,3041,2987]
a: [3,3046,1,2,1,3042]
s: 3046
i: [1,1461,0,6,3,0,0,255] cc4_struct_0
s: 1461
d: [2,3057,1,4,2,1,40,1453,2908]
a: [3,3080,1,2,1,3057]
s: 3080
i: [1,1505,0,6,2,0,0,255] t1_subset
s: 1505
i: [1,1514,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 1514
d: [2,3035,2,2,1,1,40,3000,1439]
a: [3,3039,2,2,1,3035]
s: 3039
i: [1,1446,0,7,2,0,0,1] dt_k4_struct_0
s: 1446
d: [2,3048,2,2,1,1,40,3005,1439]
a: [3,3052,2,2,1,3048]
s: 3052
i: [1,1447,0,7,2,0,0,255] rc17_struct_0
s: 1447
i: [1,1459,0,7,3,0,0,255] cc8_struct_0
s: 1459
d: [2,3055,2,4,2,1,40,1453,2975]
a: [3,3062,2,2,1,3055]
s: 3062
i: [1,1462,0,7,3,0,0,255] cc11_struct_0
s: 1462
d: [2,3056,2,4,2,1,40,1453,2982]
a: [3,3071,2,2,1,3056]
s: 3071
i: [1,1463,0,7,3,0,0,255] cc11_struct_0
s: 1463
i: [1,1464,0,7,3,0,0,255] cc9_struct_0
s: 1464
d: [2,3110,3,4,2,1,40,3039,1455]
a: [3,3115,3,2,1,3110]
s: 3115
i: [1,1466,0,7,3,0,0,255] fc9_struct_0
s: 1466
i: [1,1469,0,7,3,0,0,255] fc6_struct_0
s: 1469
i: [1,1471,0,7,3,0,0,255] rc22_struct_0
s: 1471
i: [1,1496,0,7,3,0,0,2] fc2_struct_0
s: 1496
i: [1,1498,0,7,3,0,0,2] rc4_struct_0
s: 1498
i: [1,1500,0,7,3,0,0,255] rc21_struct_0
s: 1500
i: [1,1501,0,7,3,0,0,255] rc21_struct_0
s: 1501
i: [1,1502,0,7,3,0,0,255] fc8_struct_0
s: 1502
i: [1,1503,0,7,3,0,0,255] fc7_struct_0
s: 1503
i: [1,1504,0,7,3,0,0,2] fc1_struct_0
s: 1504
i: [1,1507,0,7,2,0,0,255] t3_subset
s: 1507
i: [1,1508,0,7,3,0,0,1] t8_boole
s: 1508
i: [1,1468,0,8,3,0,0,255] rc19_struct_0
s: 1468
i: [1,1506,0,8,3,0,0,255] t2_subset
s: 1506
r: 1575
r: 1580
r: 1689
r: 1773
a: [3,1574,0,4,1,1433]
s: 1574
a: [3,1579,0,4,1,1432]
s: 1579
i: [1,1427,0,95,7,1,0,0] t44_bciideal
a: [3,1593,0,4,1,1427]
s: 1593
r: 1593
a: [3,1592,0,4,1,1427]
s: 1592
i: [1,1423,0,95,7,1,0,0] t44_bciideal
a: [3,1601,0,4,1,1423]
s: 1601
r: 1601
a: [3,1600,0,4,1,1423]
s: 1600
i: [1,1419,0,86,7,1,0,0] t44_bciideal
a: [3,1609,0,4,1,1419]
s: 1609
i: [1,1403,0,86,7,1,0,0] t44_bciideal
a: [3,1629,0,4,1,1403]
s: 1629
r: 1574
r: 1579
r: 1629
s: 1575
s: 1580
d: [2,3138,1,7,2,1,40,1506,1773]
a: [3,3153,1,3,1,3138]
s: 3153
d: [2,3192,3,0,0,3,52,2873,3029,3153,1496]
e: 3192
d: [2,3193,2,4,2,1,40,3153,1496]
d: [2,3195,2,2,1,2,21,3193,3029]
d: [2,3196,2,0,0,3,21,3195,2873]
e: 3196
r: 1575
r: 1580
r: 3153
s: 1574
a: [3,3152,1,3,1,3138]
s: 3152
s: 1579
s: 1629
i: [1,1387,0,86,7,1,0,0] t44_bciideal
a: [3,1649,0,4,1,1387]
s: 1649
i: [1,1371,0,86,7,1,0,0] t44_bciideal
a: [3,1669,0,4,1,1371]
s: 1669
a: [3,1688,0,4,1,1355]
s: 1688
a: [3,1772,0,4,1,1275]
s: 1772
a: [3,1553,0,6,1,1435]
s: 1553
d: [2,3182,1,7,2,1,40,1629,1506]
a: [3,3186,1,4,1,3182]
s: 3186
i: [1,1431,0,97,7,1,0,0] t44_bciideal
a: [3,1585,0,6,1,1431]
s: 1585
d: [2,3202,1,7,2,1,40,1649,1506]
d: [2,3203,1,4,1,2,21,3202,3152]
a: [3,3207,1,4,1,3203]
s: 3207
a: [3,1548,0,10,1,1435]
s: 1548
a: [3,1557,0,15,1,1435]
s: 1557
d: [2,3209,1,7,2,1,40,1669,1506]
d: [2,3210,1,4,1,2,21,3209,3152]
a: [3,3214,1,4,1,3210]
s: 3214
i: [1,1465,0,9,4,0,0,255] cc10_struct_0
s: 1465
d: [2,3216,2,4,1,1,40,3186,1514]
a: [3,3222,2,4,1,3216]
s: 3222
i: [1,1467,0,9,3,0,0,255] rc19_struct_0
s: 1467
i: [1,1511,0,9,3,0,0,255] t5_subset
s: 1511
d: [2,3224,2,4,1,1,40,3207,1514]
a: [3,3230,2,4,1,3224]
s: 3230
i: [1,1470,0,10,3,0,0,255] rc22_struct_0
s: 1470
d: [2,3232,2,4,1,1,40,3214,1514]
a: [3,3238,2,4,1,3232]
s: 3238
i: [1,1497,0,10,3,0,0,2] rc4_struct_0
s: 1497
i: [1,1499,0,10,3,0,0,255] rc21_struct_0
s: 1499
d: [2,3144,1,6,2,0,40,1506,1512]
s: 3144
i: [1,1510,0,10,3,0,0,255] t4_subset
s: 1510
d: [2,3253,2,6,2,0,40,3144,1514]
s: 3253
i: [1,1441,0,12,2,0,0,255] dt_u1_bcialg_1
s: 1441
i: [1,1442,0,14,2,0,0,255] dt_u1_bcialg_1
s: 1442
d: [2,3033,1,7,3,0,44,1449,1449]
s: 3033
i: [1,1494,0,16,7,0,0,1] existence_m2_bcialg_1
s: 1494
d: [2,3246,1,7,2,0,40,1511,1512]
s: 3246
i: [1,1490,0,17,8,0,0,1] rc2_bciideal
s: 1490
i: [1,1492,0,17,8,0,0,1] dt_m2_bcialg_1
s: 1492
d: [2,3265,1,16,7,1,40,1492,2833]
d: [2,3268,1,14,6,2,21,3265,2843]
d: [2,3269,1,12,5,3,21,3268,2853]
d: [2,3270,1,10,4,4,21,3269,2858]
d: [2,3271,1,8,3,5,21,3270,2863]
d: [2,3272,1,6,2,6,21,3271,2868]
d: [2,3273,1,4,1,7,21,3272,2873]
a: [3,3277,1,4,1,3273]
s: 3277
i: [1,1495,0,17,7,0,0,255] rc1_bcialg_2
s: 1495
d: [2,3264,2,6,2,0,40,3246,3144]
s: 3264
i: [1,1515,0,17,4,0,0,255] dt_k1_bcialg_1
s: 1515
i: [1,1489,0,18,8,0,0,1] rc2_bciideal
s: 1489
d: [2,3245,1,8,3,0,40,1511,1507]
s: 3245
i: [1,1491,0,18,8,0,0,1] rc2_bciideal
s: 1491
d: [2,3260,1,8,2,0,40,1510,1512]
s: 3260
i: [1,1493,0,20,8,0,0,1] dt_m2_bcialg_1
s: 1493
i: [1,1487,0,35,16,0,0,1] t36_bciideal
d: [2,1533,0,21,9,0,18,1487]
s: 1533
d: [2,3128,1,9,4,0,44,1459,1449]
s: 3128
i: [1,1488,0,35,16,0,0,1] t36_bciideal
d: [2,1534,0,21,9,0,18,1488]
s: 1534
d: [2,3130,1,9,4,0,44,1464,1449]
s: 3130
i: [1,1443,0,22,4,0,0,255] d1_bcialg_1
s: 1443
i: [1,1485,0,26,10,0,0,1] t36_bciideal
s: 1485
d: [2,3304,1,21,9,1,40,1485,2833]
d: [2,3307,1,16,8,2,21,3304,2838]
d: [2,3308,1,14,7,3,21,3307,2843]
d: [2,3309,1,12,6,4,21,3308,2848]
d: [2,3310,1,10,5,5,21,3309,2853]
d: [2,3311,1,8,4,6,21,3310,2858]
d: [2,3312,1,6,3,7,21,3311,2863]
d: [2,3313,1,4,2,8,21,3312,2868]
d: [2,3314,1,2,1,9,21,3313,2873]
a: [3,3318,1,2,1,3314]
s: 3318
i: [1,1516,0,29,6,0,0,255] dt_k5_binop_1
s: 1516
d: [2,3256,1,9,3,0,40,1510,1507]
s: 3256
i: [1,1517,0,32,6,0,0,255] redefinition_k5_binop_1
s: 1517
i: [1,1474,0,50,19,0,0,2] t7_bcialg_3
d: [2,1539,0,36,12,0,18,1474]
s: 1539
d: [2,3142,1,10,3,0,40,1506,1446]
s: 3142
i: [1,1472,0,53,18,0,0,2] t10_bcialg_3
d: [2,1540,0,39,11,0,18,1472]
s: 1540
d: [2,3143,1,10,3,0,40,1506,1447]
s: 3143
i: [1,1475,0,54,19,0,0,1] t9_bcialg_3
d: [2,1537,0,40,12,0,18,1475]
s: 1537
i: [1,1476,0,54,19,0,0,1] t9_bcialg_3
d: [2,1538,0,40,12,0,18,1476]
s: 1538
d: [2,3145,1,10,3,0,40,1506,1507]
s: 3145
i: [1,1473,0,56,18,0,0,2] t10_bcialg_3
d: [2,1541,0,42,11,0,18,1473]
s: 1541
d: [2,3286,2,10,2,0,40,3260,3144]
s: 3286
i: [1,1477,0,61,20,0,0,1] t8_bcialg_3
d: [2,1536,0,47,13,0,18,1477]
s: 1536
d: [2,3398,2,0,0,13,52,2873,2868,2863,2858,2853,2848,3318,2843,1669,1649,1629,1609,1557,1536]
e: 3398
d: [2,3410,1,41,13,1,44,1557,1536]
d: [2,3424,1,32,12,1,19,3410]
d: [2,3433,1,28,11,2,21,3424,1609]
d: [2,3434,1,24,10,3,21,3433,1629]
d: [2,3435,1,20,9,4,21,3434,1649]
d: [2,3436,1,16,8,5,21,3435,1669]
d: [2,3437,1,14,7,6,21,3436,2843]
d: [2,3438,1,12,6,7,21,3437,3318]
d: [2,3439,1,10,5,8,21,3438,2848]
d: [2,3440,1,8,4,9,21,3439,2853]
d: [2,3441,1,6,3,10,21,3440,2858]
d: [2,3442,1,4,2,11,21,3441,2863]
d: [2,3443,1,2,1,12,21,3442,2868]
d: [2,3444,1,0,0,13,21,3443,2873]
e: 3444
r: 1649
r: 1688
r: 1772
r: 1561
r: 3207
a: [3,1648,0,4,1,1387]
s: 1648
s: 1689
s: 1773
a: [3,3149,1,4,1,3138]
s: 3149
a: [3,1560,0,15,1,1435]
s: 1560
a: [3,1565,0,18,1,1435]
s: 1565
d: [2,3599,2,0,0,11,52,2873,2868,2863,2858,2853,2848,3318,2843,1773,1689,1565,1541]
e: 3599
d: [2,3600,1,33,11,1,44,1565,1541]
d: [2,3601,1,24,10,1,19,3600]
d: [2,3603,1,20,9,2,21,3601,1689]
d: [2,3604,1,16,8,3,21,3603,1773]
d: [2,3605,1,14,7,4,21,3604,2843]
d: [2,3606,1,12,6,5,21,3605,3318]
d: [2,3607,1,10,5,6,21,3606,2848]
d: [2,3608,1,8,4,7,21,3607,2853]
d: [2,3609,1,6,3,8,21,3608,2858]
d: [2,3610,1,4,2,9,21,3609,2863]
d: [2,3611,1,2,1,10,21,3610,2868]
d: [2,3612,1,0,0,11,21,3611,2873]
e: 3612
r: 1557
r: 1565
r: 1574
r: 1579
r: 1648
r: 1545
r: 1548
s: 1575
s: 1580
s: 1649
a: [3,1544,0,6,1,1435]
s: 1544
r: 1544
r: 1580
r: 1600
r: 1629
r: 3186
s: 1579
s: 1601
a: [3,1628,0,4,1,1403]
s: 1628
i: [1,1195,0,74,7,1,0,0] t44_bciideal
a: [3,1857,0,4,1,1195]
s: 1857
i: [1,955,0,74,7,1,0,0] t44_bciideal
a: [3,2101,0,4,1,955]
s: 2101
i: [1,715,0,74,7,1,0,0] t44_bciideal
a: [3,2345,0,4,1,715]
s: 2345
i: [1,475,0,74,7,1,0,0] t44_bciideal
a: [3,2589,0,4,1,475]
s: 2589
a: [3,1549,0,10,1,1435]
s: 1549
d: [2,3139,1,7,2,1,40,1506,1689]
a: [3,3158,1,4,1,3139]
s: 3158
s: 1557
a: [3,1569,0,27,1,1435]
s: 1569
d: [2,3897,1,53,13,1,44,1569,1536]
d: [2,3898,1,32,12,1,19,3897]
d: [2,3899,1,28,11,2,21,3898,1857]
d: [2,3900,1,24,10,3,21,3899,2101]
d: [2,3901,1,20,9,4,21,3900,2589]
d: [2,3902,1,16,8,5,21,3901,2345]
d: [2,3903,1,14,7,6,21,3902,2843]
d: [2,3904,1,12,6,7,21,3903,3318]
d: [2,3905,1,10,5,8,21,3904,2848]
d: [2,3906,1,8,4,9,21,3905,2853]
d: [2,3907,1,6,3,10,21,3906,2858]
d: [2,3908,1,4,2,11,21,3907,2863]
d: [2,3909,1,2,1,12,21,3908,2868]
d: [2,3910,1,0,0,13,21,3909,2873]
e: 3910
r: 1557
r: 1575
r: 1579
r: 1592
r: 1689
r: 1773
r: 2345
i: [1,1434,0,97,7,1,0,0] t44_bciideal
i: [1,1430,0,97,7,1,0,0] t44_bciideal
i: [1,1429,0,85,6,1,0,0] t44_bciideal
i: [1,1428,0,85,6,1,0,0] t44_bciideal
i: [1,1426,0,95,7,1,0,0] t44_bciideal
i: [1,1425,0,83,6,1,0,0] t44_bciideal
i: [1,1424,0,83,6,1,0,0] t44_bciideal
i: [1,1422,0,95,7,1,0,0] t44_bciideal
i: [1,1421,0,83,6,1,0,0] t44_bciideal
i: [1,1420,0,83,6,1,0,0] t44_bciideal
i: [1,1418,0,86,7,1,0,0] t44_bciideal
i: [1,1417,0,74,6,1,0,0] t44_bciideal
i: [1,1416,0,74,6,1,0,0] t44_bciideal
i: [1,1415,0,86,7,1,0,0] t44_bciideal
i: [1,1414,0,86,7,1,0,0] t44_bciideal
i: [1,1413,0,74,6,1,0,0] t44_bciideal
i: [1,1412,0,74,6,1,0,0] t44_bciideal
i: [1,1411,0,84,7,1,0,0] t44_bciideal
i: [1,1410,0,84,7,1,0,0] t44_bciideal
i: [1,1409,0,72,6,1,0,0] t44_bciideal
i: [1,1408,0,72,6,1,0,0] t44_bciideal
i: [1,1407,0,84,7,1,0,0] t44_bciideal
i: [1,1406,0,84,7,1,0,0] t44_bciideal
i: [1,1405,0,72,6,1,0,0] t44_bciideal
i: [1,1404,0,72,6,1,0,0] t44_bciideal
i: [1,1402,0,86,7,1,0,0] t44_bciideal
i: [1,1401,0,74,6,1,0,0] t44_bciideal
i: [1,1400,0,74,6,1,0,0] t44_bciideal
i: [1,1399,0,86,7,1,0,0] t44_bciideal
i: [1,1398,0,86,7,1,0,0] t44_bciideal
i: [1,1397,0,74,6,1,0,0] t44_bciideal
i: [1,1396,0,74,6,1,0,0] t44_bciideal
i: [1,1395,0,84,7,1,0,0] t44_bciideal
i: [1,1394,0,84,7,1,0,0] t44_bciideal
i: [1,1393,0,72,6,1,0,0] t44_bciideal
i: [1,1392,0,72,6,1,0,0] t44_bciideal
i: [1,1391,0,84,7,1,0,0] t44_bciideal
i: [1,1390,0,84,7,1,0,0] t44_bciideal
i: [1,1389,0,72,6,1,0,0] t44_bciideal
i: [1,1388,0,72,6,1,0,0] t44_bciideal
i: [1,1386,0,86,7,1,0,0] t44_bciideal
i: [1,1385,0,74,6,1,0,0] t44_bciideal
i: [1,1384,0,74,6,1,0,0] t44_bciideal
i: [1,1383,0,86,7,1,0,0] t44_bciideal
i: [1,1382,0,86,7,1,0,0] t44_bciideal
i: [1,1381,0,74,6,1,0,0] t44_bciideal
i: [1,1380,0,74,6,1,0,0] t44_bciideal
i: [1,1379,0,84,7,1,0,0] t44_bciideal
i: [1,1378,0,84,7,1,0,0] t44_bciideal
i: [1,1377,0,72,6,1,0,0] t44_bciideal
i: [1,1376,0,72,6,1,0,0] t44_bciideal
i: [1,1375,0,84,7,1,0,0] t44_bciideal
i: [1,1374,0,84,7,1,0,0] t44_bciideal
i: [1,1373,0,72,6,1,0,0] t44_bciideal
i: [1,1372,0,72,6,1,0,0] t44_bciideal
i: [1,1370,0,86,7,1,0,0] t44_bciideal
i: [1,1369,0,74,6,1,0,0] t44_bciideal
i: [1,1368,0,74,6,1,0,0] t44_bciideal
i: [1,1367,0,86,7,1,0,0] t44_bciideal
i: [1,1366,0,86,7,1,0,0] t44_bciideal
i: [1,1365,0,74,6,1,0,0] t44_bciideal
i: [1,1364,0,74,6,1,0,0] t44_bciideal
i: [1,1363,0,84,7,1,0,0] t44_bciideal
i: [1,1362,0,84,7,1,0,0] t44_bciideal
i: [1,1361,0,72,6,1,0,0] t44_bciideal
i: [1,1360,0,72,6,1,0,0] t44_bciideal
i: [1,1359,0,84,7,1,0,0] t44_bciideal
i: [1,1358,0,84,7,1,0,0] t44_bciideal
i: [1,1357,0,72,6,1,0,0] t44_bciideal
i: [1,1356,0,72,6,1,0,0] t44_bciideal
i: [1,1354,0,68,6,1,0,0] t44_bciideal
i: [1,1353,0,56,5,1,0,0] t44_bciideal
i: [1,1352,0,56,5,1,0,0] t44_bciideal
i: [1,1351,0,68,6,1,0,0] t44_bciideal
i: [1,1350,0,68,6,1,0,0] t44_bciideal
i: [1,1349,0,56,5,1,0,0] t44_bciideal
i: [1,1348,0,56,5,1,0,0] t44_bciideal
i: [1,1347,0,66,6,1,0,0] t44_bciideal
i: [1,1346,0,66,6,1,0,0] t44_bciideal
i: [1,1345,0,54,5,1,0,0] t44_bciideal
i: [1,1344,0,54,5,1,0,0] t44_bciideal
i: [1,1343,0,66,6,1,0,0] t44_bciideal
i: [1,1342,0,66,6,1,0,0] t44_bciideal
i: [1,1341,0,54,5,1,0,0] t44_bciideal
i: [1,1340,0,54,5,1,0,0] t44_bciideal
i: [1,1339,0,57,6,1,0,0] t44_bciideal
i: [1,1338,0,57,6,1,0,0] t44_bciideal
i: [1,1337,0,45,5,1,0,0] t44_bciideal
i: [1,1336,0,45,5,1,0,0] t44_bciideal
i: [1,1335,0,57,6,1,0,0] t44_bciideal
i: [1,1334,0,57,6,1,0,0] t44_bciideal
i: [1,1333,0,45,5,1,0,0] t44_bciideal
i: [1,1332,0,45,5,1,0,0] t44_bciideal
i: [1,1331,0,55,6,1,0,0] t44_bciideal
i: [1,1330,0,55,6,1,0,0] t44_bciideal
i: [1,1329,0,43,5,1,0,0] t44_bciideal
i: [1,1328,0,43,5,1,0,0] t44_bciideal
i: [1,1327,0,55,6,1,0,0] t44_bciideal
i: [1,1326,0,55,6,1,0,0] t44_bciideal
i: [1,1325,0,43,5,1,0,0] t44_bciideal
i: [1,1324,0,43,5,1,0,0] t44_bciideal
i: [1,1323,0,57,6,1,0,0] t44_bciideal
i: [1,1322,0,57,6,1,0,0] t44_bciideal
i: [1,1321,0,45,5,1,0,0] t44_bciideal
i: [1,1320,0,45,5,1,0,0] t44_bciideal
i: [1,1319,0,57,6,1,0,0] t44_bciideal
i: [1,1318,0,57,6,1,0,0] t44_bciideal
i: [1,1317,0,45,5,1,0,0] t44_bciideal
i: [1,1316,0,45,5,1,0,0] t44_bciideal
i: [1,1315,0,55,6,1,0,0] t44_bciideal
i: [1,1314,0,55,6,1,0,0] t44_bciideal
i: [1,1313,0,43,5,1,0,0] t44_bciideal
i: [1,1312,0,43,5,1,0,0] t44_bciideal
i: [1,1311,0,55,6,1,0,0] t44_bciideal
i: [1,1310,0,55,6,1,0,0] t44_bciideal
i: [1,1309,0,43,5,1,0,0] t44_bciideal
i: [1,1308,0,43,5,1,0,0] t44_bciideal
i: [1,1307,0,57,6,1,0,0] t44_bciideal
i: [1,1306,0,57,6,1,0,0] t44_bciideal
i: [1,1305,0,45,5,1,0,0] t44_bciideal
i: [1,1304,0,45,5,1,0,0] t44_bciideal
i: [1,1303,0,57,6,1,0,0] t44_bciideal
i: [1,1302,0,57,6,1,0,0] t44_bciideal
i: [1,1301,0,45,5,1,0,0] t44_bciideal
i: [1,1300,0,45,5,1,0,0] t44_bciideal
i: [1,1299,0,55,6,1,0,0] t44_bciideal
i: [1,1298,0,55,6,1,0,0] t44_bciideal
i: [1,1297,0,43,5,1,0,0] t44_bciideal
i: [1,1296,0,43,5,1,0,0] t44_bciideal
i: [1,1295,0,55,6,1,0,0] t44_bciideal
i: [1,1294,0,55,6,1,0,0] t44_bciideal
i: [1,1293,0,43,5,1,0,0] t44_bciideal
i: [1,1292,0,43,5,1,0,0] t44_bciideal
i: [1,1291,0,57,6,1,0,0] t44_bciideal
i: [1,1290,0,57,6,1,0,0] t44_bciideal
i: [1,1289,0,45,5,1,0,0] t44_bciideal
i: [1,1288,0,45,5,1,0,0] t44_bciideal
i: [1,1287,0,57,6,1,0,0] t44_bciideal
i: [1,1286,0,57,6,1,0,0] t44_bciideal
i: [1,1285,0,45,5,1,0,0] t44_bciideal
i: [1,1284,0,45,5,1,0,0] t44_bciideal
i: [1,1283,0,55,6,1,0,0] t44_bciideal
i: [1,1282,0,55,6,1,0,0] t44_bciideal
i: [1,1281,0,43,5,1,0,0] t44_bciideal
i: [1,1280,0,43,5,1,0,0] t44_bciideal
i: [1,1279,0,55,6,1,0,0] t44_bciideal
i: [1,1278,0,55,6,1,0,0] t44_bciideal
i: [1,1277,0,43,5,1,0,0] t44_bciideal
i: [1,1276,0,43,5,1,0,0] t44_bciideal
i: [1,1274,0,68,6,1,0,0] t44_bciideal
i: [1,1273,0,56,5,1,0,0] t44_bciideal
i: [1,1272,0,56,5,1,0,0] t44_bciideal
i: [1,1271,0,68,6,1,0,0] t44_bciideal
i: [1,1270,0,68,6,1,0,0] t44_bciideal
i: [1,1269,0,56,5,1,0,0] t44_bciideal
i: [1,1268,0,56,5,1,0,0] t44_bciideal
i: [1,1267,0,66,6,1,0,0] t44_bciideal
i: [1,1266,0,66,6,1,0,0] t44_bciideal
i: [1,1265,0,54,5,1,0,0] t44_bciideal
i: [1,1264,0,54,5,1,0,0] t44_bciideal
i: [1,1263,0,66,6,1,0,0] t44_bciideal
i: [1,1262,0,66,6,1,0,0] t44_bciideal
i: [1,1261,0,54,5,1,0,0] t44_bciideal
i: [1,1260,0,54,5,1,0,0] t44_bciideal
i: [1,1259,0,57,6,1,0,0] t44_bciideal
i: [1,1258,0,57,6,1,0,0] t44_bciideal
i: [1,1257,0,45,5,1,0,0] t44_bciideal
i: [1,1256,0,45,5,1,0,0] t44_bciideal
i: [1,1255,0,57,6,1,0,0] t44_bciideal
i: [1,1254,0,57,6,1,0,0] t44_bciideal
i: [1,1253,0,45,5,1,0,0] t44_bciideal
i: [1,1252,0,45,5,1,0,0] t44_bciideal
i: [1,1251,0,55,6,1,0,0] t44_bciideal
i: [1,1250,0,55,6,1,0,0] t44_bciideal
i: [1,1249,0,43,5,1,0,0] t44_bciideal
i: [1,1248,0,43,5,1,0,0] t44_bciideal
i: [1,1247,0,55,6,1,0,0] t44_bciideal
i: [1,1246,0,55,6,1,0,0] t44_bciideal
i: [1,1245,0,43,5,1,0,0] t44_bciideal
i: [1,1244,0,43,5,1,0,0] t44_bciideal
i: [1,1243,0,57,6,1,0,0] t44_bciideal
i: [1,1242,0,57,6,1,0,0] t44_bciideal
i: [1,1241,0,45,5,1,0,0] t44_bciideal
i: [1,1240,0,45,5,1,0,0] t44_bciideal
i: [1,1239,0,57,6,1,0,0] t44_bciideal
i: [1,1238,0,57,6,1,0,0] t44_bciideal
i: [1,1237,0,45,5,1,0,0] t44_bciideal
i: [1,1236,0,45,5,1,0,0] t44_bciideal
i: [1,1235,0,55,6,1,0,0] t44_bciideal
i: [1,1234,0,55,6,1,0,0] t44_bciideal
i: [1,1233,0,43,5,1,0,0] t44_bciideal
i: [1,1232,0,43,5,1,0,0] t44_bciideal
i: [1,1231,0,55,6,1,0,0] t44_bciideal
i: [1,1230,0,55,6,1,0,0] t44_bciideal
i: [1,1229,0,43,5,1,0,0] t44_bciideal
i: [1,1228,0,43,5,1,0,0] t44_bciideal
i: [1,1227,0,57,6,1,0,0] t44_bciideal
i: [1,1226,0,57,6,1,0,0] t44_bciideal
i: [1,1225,0,45,5,1,0,0] t44_bciideal
i: [1,1224,0,45,5,1,0,0] t44_bciideal
i: [1,1223,0,57,6,1,0,0] t44_bciideal
i: [1,1222,0,57,6,1,0,0] t44_bciideal
i: [1,1221,0,45,5,1,0,0] t44_bciideal
i: [1,1220,0,45,5,1,0,0] t44_bciideal
i: [1,1219,0,55,6,1,0,0] t44_bciideal
i: [1,1218,0,55,6,1,0,0] t44_bciideal
i: [1,1217,0,43,5,1,0,0] t44_bciideal
i: [1,1216,0,43,5,1,0,0] t44_bciideal
i: [1,1215,0,55,6,1,0,0] t44_bciideal
i: [1,1214,0,55,6,1,0,0] t44_bciideal
i: [1,1213,0,43,5,1,0,0] t44_bciideal
i: [1,1212,0,43,5,1,0,0] t44_bciideal
i: [1,1211,0,57,6,1,0,0] t44_bciideal
i: [1,1210,0,57,6,1,0,0] t44_bciideal
i: [1,1209,0,45,5,1,0,0] t44_bciideal
i: [1,1208,0,45,5,1,0,0] t44_bciideal
i: [1,1207,0,57,6,1,0,0] t44_bciideal
i: [1,1206,0,57,6,1,0,0] t44_bciideal
i: [1,1205,0,45,5,1,0,0] t44_bciideal
i: [1,1204,0,45,5,1,0,0] t44_bciideal
i: [1,1203,0,55,6,1,0,0] t44_bciideal
i: [1,1202,0,55,6,1,0,0] t44_bciideal
i: [1,1201,0,43,5,1,0,0] t44_bciideal
i: [1,1200,0,43,5,1,0,0] t44_bciideal
i: [1,1199,0,55,6,1,0,0] t44_bciideal
i: [1,1198,0,55,6,1,0,0] t44_bciideal
i: [1,1197,0,43,5,1,0,0] t44_bciideal
i: [1,1196,0,43,5,1,0,0] t44_bciideal
i: [1,1194,0,74,7,1,0,0] t44_bciideal
i: [1,1193,0,62,6,1,0,0] t44_bciideal
i: [1,1192,0,62,6,1,0,0] t44_bciideal
i: [1,1191,0,74,7,1,0,0] t44_bciideal
i: [1,1190,0,74,7,1,0,0] t44_bciideal
i: [1,1189,0,62,6,1,0,0] t44_bciideal
i: [1,1188,0,62,6,1,0,0] t44_bciideal
i: [1,1187,0,72,7,1,0,0] t44_bciideal
i: [1,1186,0,72,7,1,0,0] t44_bciideal
i: [1,1185,0,60,6,1,0,0] t44_bciideal
i: [1,1184,0,60,6,1,0,0] t44_bciideal
i: [1,1183,0,72,7,1,0,0] t44_bciideal
i: [1,1182,0,72,7,1,0,0] t44_bciideal
i: [1,1181,0,60,6,1,0,0] t44_bciideal
i: [1,1180,0,60,6,1,0,0] t44_bciideal
i: [1,1179,0,63,7,1,0,0] t44_bciideal
i: [1,1178,0,63,7,1,0,0] t44_bciideal
i: [1,1177,0,51,6,1,0,0] t44_bciideal
i: [1,1176,0,51,6,1,0,0] t44_bciideal
i: [1,1175,0,63,7,1,0,0] t44_bciideal
i: [1,1174,0,63,7,1,0,0] t44_bciideal
i: [1,1173,0,51,6,1,0,0] t44_bciideal
i: [1,1172,0,51,6,1,0,0] t44_bciideal
i: [1,1171,0,61,7,1,0,0] t44_bciideal
i: [1,1170,0,61,7,1,0,0] t44_bciideal
i: [1,1169,0,49,6,1,0,0] t44_bciideal
i: [1,1168,0,49,6,1,0,0] t44_bciideal
i: [1,1167,0,61,7,1,0,0] t44_bciideal
i: [1,1166,0,61,7,1,0,0] t44_bciideal
i: [1,1165,0,49,6,1,0,0] t44_bciideal
i: [1,1164,0,49,6,1,0,0] t44_bciideal
i: [1,1163,0,63,7,1,0,0] t44_bciideal
i: [1,1162,0,63,7,1,0,0] t44_bciideal
i: [1,1161,0,51,6,1,0,0] t44_bciideal
i: [1,1160,0,51,6,1,0,0] t44_bciideal
i: [1,1159,0,63,7,1,0,0] t44_bciideal
i: [1,1158,0,63,7,1,0,0] t44_bciideal
i: [1,1157,0,51,6,1,0,0] t44_bciideal
i: [1,1156,0,51,6,1,0,0] t44_bciideal
i: [1,1155,0,61,7,1,0,0] t44_bciideal
i: [1,1154,0,61,7,1,0,0] t44_bciideal
i: [1,1153,0,49,6,1,0,0] t44_bciideal
i: [1,1152,0,49,6,1,0,0] t44_bciideal
i: [1,1151,0,61,7,1,0,0] t44_bciideal
i: [1,1150,0,61,7,1,0,0] t44_bciideal
i: [1,1149,0,49,6,1,0,0] t44_bciideal
i: [1,1148,0,49,6,1,0,0] t44_bciideal
i: [1,1147,0,63,7,1,0,0] t44_bciideal
i: [1,1146,0,63,7,1,0,0] t44_bciideal
i: [1,1145,0,51,6,1,0,0] t44_bciideal
i: [1,1144,0,51,6,1,0,0] t44_bciideal
i: [1,1143,0,63,7,1,0,0] t44_bciideal
i: [1,1142,0,63,7,1,0,0] t44_bciideal
i: [1,1141,0,51,6,1,0,0] t44_bciideal
i: [1,1140,0,51,6,1,0,0] t44_bciideal
i: [1,1139,0,61,7,1,0,0] t44_bciideal
i: [1,1138,0,61,7,1,0,0] t44_bciideal
i: [1,1137,0,49,6,1,0,0] t44_bciideal
i: [1,1136,0,49,6,1,0,0] t44_bciideal
i: [1,1135,0,61,7,1,0,0] t44_bciideal
i: [1,1134,0,61,7,1,0,0] t44_bciideal
i: [1,1133,0,49,6,1,0,0] t44_bciideal
i: [1,1132,0,49,6,1,0,0] t44_bciideal
i: [1,1131,0,63,7,1,0,0] t44_bciideal
i: [1,1130,0,63,7,1,0,0] t44_bciideal
i: [1,1129,0,51,6,1,0,0] t44_bciideal
i: [1,1128,0,51,6,1,0,0] t44_bciideal
i: [1,1127,0,63,7,1,0,0] t44_bciideal
i: [1,1126,0,63,7,1,0,0] t44_bciideal
i: [1,1125,0,51,6,1,0,0] t44_bciideal
i: [1,1124,0,51,6,1,0,0] t44_bciideal
i: [1,1123,0,61,7,1,0,0] t44_bciideal
i: [1,1122,0,61,7,1,0,0] t44_bciideal
i: [1,1121,0,49,6,1,0,0] t44_bciideal
i: [1,1120,0,49,6,1,0,0] t44_bciideal
i: [1,1119,0,61,7,1,0,0] t44_bciideal
i: [1,1118,0,61,7,1,0,0] t44_bciideal
i: [1,1117,0,49,6,1,0,0] t44_bciideal
i: [1,1116,0,49,6,1,0,0] t44_bciideal
i: [1,1115,0,45,6,1,0,0] t44_bciideal
i: [1,1114,0,45,6,1,0,0] t44_bciideal
i: [1,1113,0,33,5,1,0,0] t44_bciideal
i: [1,1112,0,33,5,1,0,0] t44_bciideal
i: [1,1111,0,45,6,1,0,0] t44_bciideal
i: [1,1110,0,45,6,1,0,0] t44_bciideal
i: [1,1109,0,33,5,1,0,0] t44_bciideal
i: [1,1108,0,33,5,1,0,0] t44_bciideal
i: [1,1107,0,43,6,1,0,0] t44_bciideal
i: [1,1106,0,43,6,1,0,0] t44_bciideal
i: [1,1105,0,31,5,1,0,0] t44_bciideal
i: [1,1104,0,31,5,1,0,0] t44_bciideal
i: [1,1103,0,43,6,1,0,0] t44_bciideal
i: [1,1102,0,43,6,1,0,0] t44_bciideal
i: [1,1101,0,31,5,1,0,0] t44_bciideal
i: [1,1100,0,31,5,1,0,0] t44_bciideal
i: [1,1099,0,34,6,1,0,0] t44_bciideal
i: [1,1098,0,34,6,1,0,0] t44_bciideal
i: [1,1097,0,22,5,1,0,0] t44_bciideal
i: [1,1096,0,22,5,1,0,0] t44_bciideal
i: [1,1095,0,34,6,1,0,0] t44_bciideal
i: [1,1094,0,34,6,1,0,0] t44_bciideal
i: [1,1093,0,22,5,1,0,0] t44_bciideal
i: [1,1092,0,22,5,1,0,0] t44_bciideal
i: [1,1091,0,32,6,1,0,0] t44_bciideal
i: [1,1090,0,32,6,1,0,0] t44_bciideal
i: [1,1089,0,20,5,1,0,0] t44_bciideal
i: [1,1088,0,20,5,1,0,0] t44_bciideal
i: [1,1087,0,32,6,1,0,0] t44_bciideal
i: [1,1086,0,32,6,1,0,0] t44_bciideal
i: [1,1085,0,20,5,1,0,0] t44_bciideal
i: [1,1084,0,20,5,1,0,0] t44_bciideal
i: [1,1083,0,34,6,1,0,0] t44_bciideal
i: [1,1082,0,34,6,1,0,0] t44_bciideal
i: [1,1081,0,22,5,1,0,0] t44_bciideal
i: [1,1080,0,22,5,1,0,0] t44_bciideal
i: [1,1079,0,34,6,1,0,0] t44_bciideal
i: [1,1078,0,34,6,1,0,0] t44_bciideal
i: [1,1077,0,22,5,1,0,0] t44_bciideal
i: [1,1076,0,22,5,1,0,0] t44_bciideal
i: [1,1075,0,32,6,1,0,0] t44_bciideal
i: [1,1074,0,32,6,1,0,0] t44_bciideal
i: [1,1073,0,20,5,1,0,0] t44_bciideal
i: [1,1072,0,20,5,1,0,0] t44_bciideal
i: [1,1071,0,32,6,1,0,0] t44_bciideal
i: [1,1070,0,32,6,1,0,0] t44_bciideal
i: [1,1069,0,20,5,1,0,0] t44_bciideal
i: [1,1068,0,20,5,1,0,0] t44_bciideal
i: [1,1067,0,34,6,1,0,0] t44_bciideal
i: [1,1066,0,34,6,1,0,0] t44_bciideal
i: [1,1065,0,22,5,1,0,0] t44_bciideal
i: [1,1064,0,22,5,1,0,0] t44_bciideal
i: [1,1063,0,34,6,1,0,0] t44_bciideal
i: [1,1062,0,34,6,1,0,0] t44_bciideal
i: [1,1061,0,22,5,1,0,0] t44_bciideal
i: [1,1060,0,22,5,1,0,0] t44_bciideal
i: [1,1059,0,32,6,1,0,0] t44_bciideal
i: [1,1058,0,32,6,1,0,0] t44_bciideal
i: [1,1057,0,20,5,1,0,0] t44_bciideal
i: [1,1056,0,20,5,1,0,0] t44_bciideal
i: [1,1055,0,32,6,1,0,0] t44_bciideal
i: [1,1054,0,32,6,1,0,0] t44_bciideal
i: [1,1053,0,20,5,1,0,0] t44_bciideal
i: [1,1052,0,20,5,1,0,0] t44_bciideal
i: [1,1051,0,34,6,1,0,0] t44_bciideal
i: [1,1050,0,34,6,1,0,0] t44_bciideal
i: [1,1049,0,22,5,1,0,0] t44_bciideal
i: [1,1048,0,22,5,1,0,0] t44_bciideal
i: [1,1047,0,34,6,1,0,0] t44_bciideal
i: [1,1046,0,34,6,1,0,0] t44_bciideal
i: [1,1045,0,22,5,1,0,0] t44_bciideal
i: [1,1044,0,22,5,1,0,0] t44_bciideal
i: [1,1043,0,32,6,1,0,0] t44_bciideal
i: [1,1042,0,32,6,1,0,0] t44_bciideal
i: [1,1041,0,20,5,1,0,0] t44_bciideal
i: [1,1040,0,20,5,1,0,0] t44_bciideal
i: [1,1039,0,32,6,1,0,0] t44_bciideal
i: [1,1038,0,32,6,1,0,0] t44_bciideal
i: [1,1037,0,20,5,1,0,0] t44_bciideal
i: [1,1036,0,20,5,1,0,0] t44_bciideal
i: [1,1035,0,45,6,1,0,0] t44_bciideal
i: [1,1034,0,45,6,1,0,0] t44_bciideal
i: [1,1033,0,33,5,1,0,0] t44_bciideal
i: [1,1032,0,33,5,1,0,0] t44_bciideal
i: [1,1031,0,45,6,1,0,0] t44_bciideal
i: [1,1030,0,45,6,1,0,0] t44_bciideal
i: [1,1029,0,33,5,1,0,0] t44_bciideal
i: [1,1028,0,33,5,1,0,0] t44_bciideal
i: [1,1027,0,43,6,1,0,0] t44_bciideal
i: [1,1026,0,43,6,1,0,0] t44_bciideal
i: [1,1025,0,31,5,1,0,0] t44_bciideal
i: [1,1024,0,31,5,1,0,0] t44_bciideal
i: [1,1023,0,43,6,1,0,0] t44_bciideal
i: [1,1022,0,43,6,1,0,0] t44_bciideal
i: [1,1021,0,31,5,1,0,0] t44_bciideal
i: [1,1020,0,31,5,1,0,0] t44_bciideal
i: [1,1019,0,34,6,1,0,0] t44_bciideal
i: [1,1018,0,34,6,1,0,0] t44_bciideal
i: [1,1017,0,22,5,1,0,0] t44_bciideal
i: [1,1016,0,22,5,1,0,0] t44_bciideal
i: [1,1015,0,34,6,1,0,0] t44_bciideal
i: [1,1014,0,34,6,1,0,0] t44_bciideal
i: [1,1013,0,22,5,1,0,0] t44_bciideal
i: [1,1012,0,22,5,1,0,0] t44_bciideal
i: [1,1011,0,32,6,1,0,0] t44_bciideal
i: [1,1010,0,32,6,1,0,0] t44_bciideal
i: [1,1009,0,20,5,1,0,0] t44_bciideal
i: [1,1008,0,20,5,1,0,0] t44_bciideal
i: [1,1007,0,32,6,1,0,0] t44_bciideal
i: [1,1006,0,32,6,1,0,0] t44_bciideal
i: [1,1005,0,20,5,1,0,0] t44_bciideal
i: [1,1004,0,20,5,1,0,0] t44_bciideal
i: [1,1003,0,34,6,1,0,0] t44_bciideal
i: [1,1002,0,34,6,1,0,0] t44_bciideal
i: [1,1001,0,22,5,1,0,0] t44_bciideal
i: [1,1000,0,22,5,1,0,0] t44_bciideal
i: [1,999,0,34,6,1,0,0] t44_bciideal
i: [1,998,0,34,6,1,0,0] t44_bciideal
i: [1,997,0,22,5,1,0,0] t44_bciideal
i: [1,996,0,22,5,1,0,0] t44_bciideal
i: [1,995,0,32,6,1,0,0] t44_bciideal
i: [1,994,0,32,6,1,0,0] t44_bciideal
i: [1,993,0,20,5,1,0,0] t44_bciideal
i: [1,992,0,20,5,1,0,0] t44_bciideal
i: [1,991,0,32,6,1,0,0] t44_bciideal
i: [1,990,0,32,6,1,0,0] t44_bciideal
i: [1,989,0,20,5,1,0,0] t44_bciideal
i: [1,988,0,20,5,1,0,0] t44_bciideal
i: [1,987,0,34,6,1,0,0] t44_bciideal
i: [1,986,0,34,6,1,0,0] t44_bciideal
i: [1,985,0,22,5,1,0,0] t44_bciideal
i: [1,984,0,22,5,1,0,0] t44_bciideal
i: [1,983,0,34,6,1,0,0] t44_bciideal
i: [1,982,0,34,6,1,0,0] t44_bciideal
i: [1,981,0,22,5,1,0,0] t44_bciideal
i: [1,980,0,22,5,1,0,0] t44_bciideal
i: [1,979,0,32,6,1,0,0] t44_bciideal
i: [1,978,0,32,6,1,0,0] t44_bciideal
i: [1,977,0,20,5,1,0,0] t44_bciideal
i: [1,976,0,20,5,1,0,0] t44_bciideal
i: [1,975,0,32,6,1,0,0] t44_bciideal
i: [1,974,0,32,6,1,0,0] t44_bciideal
i: [1,973,0,20,5,1,0,0] t44_bciideal
i: [1,972,0,20,5,1,0,0] t44_bciideal
i: [1,971,0,34,6,1,0,0] t44_bciideal
i: [1,970,0,34,6,1,0,0] t44_bciideal
i: [1,969,0,22,5,1,0,0] t44_bciideal
i: [1,968,0,22,5,1,0,0] t44_bciideal
i: [1,967,0,34,6,1,0,0] t44_bciideal
i: [1,966,0,34,6,1,0,0] t44_bciideal
i: [1,965,0,22,5,1,0,0] t44_bciideal
i: [1,964,0,22,5,1,0,0] t44_bciideal
i: [1,963,0,32,6,1,0,0] t44_bciideal
i: [1,962,0,32,6,1,0,0] t44_bciideal
i: [1,961,0,20,5,1,0,0] t44_bciideal
i: [1,960,0,20,5,1,0,0] t44_bciideal
i: [1,959,0,32,6,1,0,0] t44_bciideal
i: [1,958,0,32,6,1,0,0] t44_bciideal
i: [1,957,0,20,5,1,0,0] t44_bciideal
i: [1,956,0,20,5,1,0,0] t44_bciideal
i: [1,954,0,74,7,1,0,0] t44_bciideal
i: [1,953,0,62,6,1,0,0] t44_bciideal
i: [1,952,0,62,6,1,0,0] t44_bciideal
i: [1,951,0,74,7,1,0,0] t44_bciideal
i: [1,950,0,74,7,1,0,0] t44_bciideal
i: [1,949,0,62,6,1,0,0] t44_bciideal
i: [1,948,0,62,6,1,0,0] t44_bciideal
i: [1,947,0,72,7,1,0,0] t44_bciideal
i: [1,946,0,72,7,1,0,0] t44_bciideal
i: [1,945,0,60,6,1,0,0] t44_bciideal
i: [1,944,0,60,6,1,0,0] t44_bciideal
i: [1,943,0,72,7,1,0,0] t44_bciideal
i: [1,942,0,72,7,1,0,0] t44_bciideal
i: [1,941,0,60,6,1,0,0] t44_bciideal
i: [1,940,0,60,6,1,0,0] t44_bciideal
i: [1,939,0,63,7,1,0,0] t44_bciideal
i: [1,938,0,63,7,1,0,0] t44_bciideal
i: [1,937,0,51,6,1,0,0] t44_bciideal
i: [1,936,0,51,6,1,0,0] t44_bciideal
i: [1,935,0,63,7,1,0,0] t44_bciideal
i: [1,934,0,63,7,1,0,0] t44_bciideal
i: [1,933,0,51,6,1,0,0] t44_bciideal
i: [1,932,0,51,6,1,0,0] t44_bciideal
i: [1,931,0,61,7,1,0,0] t44_bciideal
i: [1,930,0,61,7,1,0,0] t44_bciideal
i: [1,929,0,49,6,1,0,0] t44_bciideal
i: [1,928,0,49,6,1,0,0] t44_bciideal
i: [1,927,0,61,7,1,0,0] t44_bciideal
i: [1,926,0,61,7,1,0,0] t44_bciideal
i: [1,925,0,49,6,1,0,0] t44_bciideal
i: [1,924,0,49,6,1,0,0] t44_bciideal
i: [1,923,0,63,7,1,0,0] t44_bciideal
i: [1,922,0,63,7,1,0,0] t44_bciideal
i: [1,921,0,51,6,1,0,0] t44_bciideal
i: [1,920,0,51,6,1,0,0] t44_bciideal
i: [1,919,0,63,7,1,0,0] t44_bciideal
i: [1,918,0,63,7,1,0,0] t44_bciideal
i: [1,917,0,51,6,1,0,0] t44_bciideal
i: [1,916,0,51,6,1,0,0] t44_bciideal
i: [1,915,0,61,7,1,0,0] t44_bciideal
i: [1,914,0,61,7,1,0,0] t44_bciideal
i: [1,913,0,49,6,1,0,0] t44_bciideal
i: [1,912,0,49,6,1,0,0] t44_bciideal
i: [1,911,0,61,7,1,0,0] t44_bciideal
i: [1,910,0,61,7,1,0,0] t44_bciideal
i: [1,909,0,49,6,1,0,0] t44_bciideal
i: [1,908,0,49,6,1,0,0] t44_bciideal
i: [1,907,0,63,7,1,0,0] t44_bciideal
i: [1,906,0,63,7,1,0,0] t44_bciideal
i: [1,905,0,51,6,1,0,0] t44_bciideal
i: [1,904,0,51,6,1,0,0] t44_bciideal
i: [1,903,0,63,7,1,0,0] t44_bciideal
i: [1,902,0,63,7,1,0,0] t44_bciideal
i: [1,901,0,51,6,1,0,0] t44_bciideal
i: [1,900,0,51,6,1,0,0] t44_bciideal
i: [1,899,0,61,7,1,0,0] t44_bciideal
i: [1,898,0,61,7,1,0,0] t44_bciideal
i: [1,897,0,49,6,1,0,0] t44_bciideal
i: [1,896,0,49,6,1,0,0] t44_bciideal
i: [1,895,0,61,7,1,0,0] t44_bciideal
i: [1,894,0,61,7,1,0,0] t44_bciideal
i: [1,893,0,49,6,1,0,0] t44_bciideal
i: [1,892,0,49,6,1,0,0] t44_bciideal
i: [1,891,0,63,7,1,0,0] t44_bciideal
i: [1,890,0,63,7,1,0,0] t44_bciideal
i: [1,889,0,51,6,1,0,0] t44_bciideal
i: [1,888,0,51,6,1,0,0] t44_bciideal
i: [1,887,0,63,7,1,0,0] t44_bciideal
i: [1,886,0,63,7,1,0,0] t44_bciideal
i: [1,885,0,51,6,1,0,0] t44_bciideal
i: [1,884,0,51,6,1,0,0] t44_bciideal
i: [1,883,0,61,7,1,0,0] t44_bciideal
i: [1,882,0,61,7,1,0,0] t44_bciideal
i: [1,881,0,49,6,1,0,0] t44_bciideal
i: [1,880,0,49,6,1,0,0] t44_bciideal
i: [1,879,0,61,7,1,0,0] t44_bciideal
i: [1,878,0,61,7,1,0,0] t44_bciideal
i: [1,877,0,49,6,1,0,0] t44_bciideal
i: [1,876,0,49,6,1,0,0] t44_bciideal
i: [1,875,0,45,6,1,0,0] t44_bciideal
i: [1,874,0,45,6,1,0,0] t44_bciideal
i: [1,873,0,33,5,1,0,0] t44_bciideal
i: [1,872,0,33,5,1,0,0] t44_bciideal
i: [1,871,0,45,6,1,0,0] t44_bciideal
i: [1,870,0,45,6,1,0,0] t44_bciideal
i: [1,869,0,33,5,1,0,0] t44_bciideal
i: [1,868,0,33,5,1,0,0] t44_bciideal
i: [1,867,0,43,6,1,0,0] t44_bciideal
i: [1,866,0,43,6,1,0,0] t44_bciideal
i: [1,865,0,31,5,1,0,0] t44_bciideal
i: [1,864,0,31,5,1,0,0] t44_bciideal
i: [1,863,0,43,6,1,0,0] t44_bciideal
i: [1,862,0,43,6,1,0,0] t44_bciideal
i: [1,861,0,31,5,1,0,0] t44_bciideal
i: [1,860,0,31,5,1,0,0] t44_bciideal
i: [1,859,0,34,6,1,0,0] t44_bciideal
i: [1,858,0,34,6,1,0,0] t44_bciideal
i: [1,857,0,22,5,1,0,0] t44_bciideal
i: [1,856,0,22,5,1,0,0] t44_bciideal
i: [1,855,0,34,6,1,0,0] t44_bciideal
i: [1,854,0,34,6,1,0,0] t44_bciideal
i: [1,853,0,22,5,1,0,0] t44_bciideal
i: [1,852,0,22,5,1,0,0] t44_bciideal
i: [1,851,0,32,6,1,0,0] t44_bciideal
i: [1,850,0,32,6,1,0,0] t44_bciideal
i: [1,849,0,20,5,1,0,0] t44_bciideal
i: [1,848,0,20,5,1,0,0] t44_bciideal
i: [1,847,0,32,6,1,0,0] t44_bciideal
i: [1,846,0,32,6,1,0,0] t44_bciideal
i: [1,845,0,20,5,1,0,0] t44_bciideal
i: [1,844,0,20,5,1,0,0] t44_bciideal
i: [1,843,0,34,6,1,0,0] t44_bciideal
i: [1,842,0,34,6,1,0,0] t44_bciideal
i: [1,841,0,22,5,1,0,0] t44_bciideal
i: [1,840,0,22,5,1,0,0] t44_bciideal
i: [1,839,0,34,6,1,0,0] t44_bciideal
i: [1,838,0,34,6,1,0,0] t44_bciideal
i: [1,837,0,22,5,1,0,0] t44_bciideal
i: [1,836,0,22,5,1,0,0] t44_bciideal
i: [1,835,0,32,6,1,0,0] t44_bciideal
i: [1,834,0,32,6,1,0,0] t44_bciideal
i: [1,833,0,20,5,1,0,0] t44_bciideal
i: [1,832,0,20,5,1,0,0] t44_bciideal
i: [1,831,0,32,6,1,0,0] t44_bciideal
i: [1,830,0,32,6,1,0,0] t44_bciideal
i: [1,829,0,20,5,1,0,0] t44_bciideal
i: [1,828,0,20,5,1,0,0] t44_bciideal
i: [1,827,0,34,6,1,0,0] t44_bciideal
i: [1,826,0,34,6,1,0,0] t44_bciideal
i: [1,825,0,22,5,1,0,0] t44_bciideal
i: [1,824,0,22,5,1,0,0] t44_bciideal
i: [1,823,0,34,6,1,0,0] t44_bciideal
i: [1,822,0,34,6,1,0,0] t44_bciideal
i: [1,821,0,22,5,1,0,0] t44_bciideal
i: [1,820,0,22,5,1,0,0] t44_bciideal
i: [1,819,0,32,6,1,0,0] t44_bciideal
i: [1,818,0,32,6,1,0,0] t44_bciideal
i: [1,817,0,20,5,1,0,0] t44_bciideal
i: [1,816,0,20,5,1,0,0] t44_bciideal
i: [1,815,0,32,6,1,0,0] t44_bciideal
i: [1,814,0,32,6,1,0,0] t44_bciideal
i: [1,813,0,20,5,1,0,0] t44_bciideal
i: [1,812,0,20,5,1,0,0] t44_bciideal
i: [1,811,0,34,6,1,0,0] t44_bciideal
i: [1,810,0,34,6,1,0,0] t44_bciideal
i: [1,809,0,22,5,1,0,0] t44_bciideal
i: [1,808,0,22,5,1,0,0] t44_bciideal
i: [1,807,0,34,6,1,0,0] t44_bciideal
i: [1,806,0,34,6,1,0,0] t44_bciideal
i: [1,805,0,22,5,1,0,0] t44_bciideal
i: [1,804,0,22,5,1,0,0] t44_bciideal
i: [1,803,0,32,6,1,0,0] t44_bciideal
i: [1,802,0,32,6,1,0,0] t44_bciideal
i: [1,801,0,20,5,1,0,0] t44_bciideal
i: [1,800,0,20,5,1,0,0] t44_bciideal
i: [1,799,0,32,6,1,0,0] t44_bciideal
i: [1,798,0,32,6,1,0,0] t44_bciideal
i: [1,797,0,20,5,1,0,0] t44_bciideal
i: [1,796,0,20,5,1,0,0] t44_bciideal
i: [1,795,0,45,6,1,0,0] t44_bciideal
i: [1,794,0,45,6,1,0,0] t44_bciideal
i: [1,793,0,33,5,1,0,0] t44_bciideal
i: [1,792,0,33,5,1,0,0] t44_bciideal
i: [1,791,0,45,6,1,0,0] t44_bciideal
i: [1,790,0,45,6,1,0,0] t44_bciideal
i: [1,789,0,33,5,1,0,0] t44_bciideal
i: [1,788,0,33,5,1,0,0] t44_bciideal
i: [1,787,0,43,6,1,0,0] t44_bciideal
i: [1,786,0,43,6,1,0,0] t44_bciideal
i: [1,785,0,31,5,1,0,0] t44_bciideal
i: [1,784,0,31,5,1,0,0] t44_bciideal
i: [1,783,0,43,6,1,0,0] t44_bciideal
i: [1,782,0,43,6,1,0,0] t44_bciideal
i: [1,781,0,31,5,1,0,0] t44_bciideal
i: [1,780,0,31,5,1,0,0] t44_bciideal
i: [1,779,0,34,6,1,0,0] t44_bciideal
i: [1,778,0,34,6,1,0,0] t44_bciideal
i: [1,777,0,22,5,1,0,0] t44_bciideal
i: [1,776,0,22,5,1,0,0] t44_bciideal
i: [1,775,0,34,6,1,0,0] t44_bciideal
i: [1,774,0,34,6,1,0,0] t44_bciideal
i: [1,773,0,22,5,1,0,0] t44_bciideal
i: [1,772,0,22,5,1,0,0] t44_bciideal
i: [1,771,0,32,6,1,0,0] t44_bciideal
i: [1,770,0,32,6,1,0,0] t44_bciideal
i: [1,769,0,20,5,1,0,0] t44_bciideal
i: [1,768,0,20,5,1,0,0] t44_bciideal
i: [1,767,0,32,6,1,0,0] t44_bciideal
i: [1,766,0,32,6,1,0,0] t44_bciideal
i: [1,765,0,20,5,1,0,0] t44_bciideal
i: [1,764,0,20,5,1,0,0] t44_bciideal
i: [1,763,0,34,6,1,0,0] t44_bciideal
i: [1,762,0,34,6,1,0,0] t44_bciideal
i: [1,761,0,22,5,1,0,0] t44_bciideal
i: [1,760,0,22,5,1,0,0] t44_bciideal
i: [1,759,0,34,6,1,0,0] t44_bciideal
i: [1,758,0,34,6,1,0,0] t44_bciideal
i: [1,757,0,22,5,1,0,0] t44_bciideal
i: [1,756,0,22,5,1,0,0] t44_bciideal
i: [1,755,0,32,6,1,0,0] t44_bciideal
i: [1,754,0,32,6,1,0,0] t44_bciideal
i: [1,753,0,20,5,1,0,0] t44_bciideal
i: [1,752,0,20,5,1,0,0] t44_bciideal
i: [1,751,0,32,6,1,0,0] t44_bciideal
i: [1,750,0,32,6,1,0,0] t44_bciideal
i: [1,749,0,20,5,1,0,0] t44_bciideal
i: [1,748,0,20,5,1,0,0] t44_bciideal
i: [1,747,0,34,6,1,0,0] t44_bciideal
i: [1,746,0,34,6,1,0,0] t44_bciideal
i: [1,745,0,22,5,1,0,0] t44_bciideal
i: [1,744,0,22,5,1,0,0] t44_bciideal
i: [1,743,0,34,6,1,0,0] t44_bciideal
i: [1,742,0,34,6,1,0,0] t44_bciideal
i: [1,741,0,22,5,1,0,0] t44_bciideal
i: [1,740,0,22,5,1,0,0] t44_bciideal
i: [1,739,0,32,6,1,0,0] t44_bciideal
i: [1,738,0,32,6,1,0,0] t44_bciideal
i: [1,737,0,20,5,1,0,0] t44_bciideal
i: [1,736,0,20,5,1,0,0] t44_bciideal
i: [1,735,0,32,6,1,0,0] t44_bciideal
i: [1,734,0,32,6,1,0,0] t44_bciideal
i: [1,733,0,20,5,1,0,0] t44_bciideal
i: [1,732,0,20,5,1,0,0] t44_bciideal
i: [1,731,0,34,6,1,0,0] t44_bciideal
i: [1,730,0,34,6,1,0,0] t44_bciideal
i: [1,729,0,22,5,1,0,0] t44_bciideal
i: [1,728,0,22,5,1,0,0] t44_bciideal
i: [1,727,0,34,6,1,0,0] t44_bciideal
i: [1,726,0,34,6,1,0,0] t44_bciideal
i: [1,725,0,22,5,1,0,0] t44_bciideal
i: [1,724,0,22,5,1,0,0] t44_bciideal
i: [1,723,0,32,6,1,0,0] t44_bciideal
i: [1,722,0,32,6,1,0,0] t44_bciideal
i: [1,721,0,20,5,1,0,0] t44_bciideal
i: [1,720,0,20,5,1,0,0] t44_bciideal
i: [1,719,0,32,6,1,0,0] t44_bciideal
i: [1,718,0,32,6,1,0,0] t44_bciideal
i: [1,717,0,20,5,1,0,0] t44_bciideal
i: [1,716,0,20,5,1,0,0] t44_bciideal
i: [1,714,0,74,7,1,0,0] t44_bciideal
i: [1,713,0,62,6,1,0,0] t44_bciideal
i: [1,712,0,62,6,1,0,0] t44_bciideal
i: [1,711,0,74,7,1,0,0] t44_bciideal
i: [1,710,0,74,7,1,0,0] t44_bciideal
i: [1,709,0,62,6,1,0,0] t44_bciideal
i: [1,708,0,62,6,1,0,0] t44_bciideal
i: [1,707,0,72,7,1,0,0] t44_bciideal
i: [1,706,0,72,7,1,0,0] t44_bciideal
i: [1,705,0,60,6,1,0,0] t44_bciideal
i: [1,704,0,60,6,1,0,0] t44_bciideal
i: [1,703,0,72,7,1,0,0] t44_bciideal
i: [1,702,0,72,7,1,0,0] t44_bciideal
i: [1,701,0,60,6,1,0,0] t44_bciideal
i: [1,700,0,60,6,1,0,0] t44_bciideal
i: [1,699,0,63,7,1,0,0] t44_bciideal
i: [1,698,0,63,7,1,0,0] t44_bciideal
i: [1,697,0,51,6,1,0,0] t44_bciideal
i: [1,696,0,51,6,1,0,0] t44_bciideal
i: [1,695,0,63,7,1,0,0] t44_bciideal
i: [1,694,0,63,7,1,0,0] t44_bciideal
i: [1,693,0,51,6,1,0,0] t44_bciideal
i: [1,692,0,51,6,1,0,0] t44_bciideal
i: [1,691,0,61,7,1,0,0] t44_bciideal
i: [1,690,0,61,7,1,0,0] t44_bciideal
i: [1,689,0,49,6,1,0,0] t44_bciideal
i: [1,688,0,49,6,1,0,0] t44_bciideal
i: [1,687,0,61,7,1,0,0] t44_bciideal
i: [1,686,0,61,7,1,0,0] t44_bciideal
i: [1,685,0,49,6,1,0,0] t44_bciideal
i: [1,684,0,49,6,1,0,0] t44_bciideal
i: [1,683,0,63,7,1,0,0] t44_bciideal
i: [1,682,0,63,7,1,0,0] t44_bciideal
i: [1,681,0,51,6,1,0,0] t44_bciideal
i: [1,680,0,51,6,1,0,0] t44_bciideal
i: [1,679,0,63,7,1,0,0] t44_bciideal
i: [1,678,0,63,7,1,0,0] t44_bciideal
i: [1,677,0,51,6,1,0,0] t44_bciideal
i: [1,676,0,51,6,1,0,0] t44_bciideal
i: [1,675,0,61,7,1,0,0] t44_bciideal
i: [1,674,0,61,7,1,0,0] t44_bciideal
i: [1,673,0,49,6,1,0,0] t44_bciideal
i: [1,672,0,49,6,1,0,0] t44_bciideal
i: [1,671,0,61,7,1,0,0] t44_bciideal
i: [1,670,0,61,7,1,0,0] t44_bciideal
i: [1,669,0,49,6,1,0,0] t44_bciideal
i: [1,668,0,49,6,1,0,0] t44_bciideal
i: [1,667,0,63,7,1,0,0] t44_bciideal
i: [1,666,0,63,7,1,0,0] t44_bciideal
i: [1,665,0,51,6,1,0,0] t44_bciideal
i: [1,664,0,51,6,1,0,0] t44_bciideal
i: [1,663,0,63,7,1,0,0] t44_bciideal
i: [1,662,0,63,7,1,0,0] t44_bciideal
i: [1,661,0,51,6,1,0,0] t44_bciideal
i: [1,660,0,51,6,1,0,0] t44_bciideal
i: [1,659,0,61,7,1,0,0] t44_bciideal
i: [1,658,0,61,7,1,0,0] t44_bciideal
i: [1,657,0,49,6,1,0,0] t44_bciideal
i: [1,656,0,49,6,1,0,0] t44_bciideal
i: [1,655,0,61,7,1,0,0] t44_bciideal
i: [1,654,0,61,7,1,0,0] t44_bciideal
i: [1,653,0,49,6,1,0,0] t44_bciideal
i: [1,652,0,49,6,1,0,0] t44_bciideal
i: [1,651,0,63,7,1,0,0] t44_bciideal
i: [1,650,0,63,7,1,0,0] t44_bciideal
i: [1,649,0,51,6,1,0,0] t44_bciideal
i: [1,648,0,51,6,1,0,0] t44_bciideal
i: [1,647,0,63,7,1,0,0] t44_bciideal
i: [1,646,0,63,7,1,0,0] t44_bciideal
i: [1,645,0,51,6,1,0,0] t44_bciideal
i: [1,644,0,51,6,1,0,0] t44_bciideal
i: [1,643,0,61,7,1,0,0] t44_bciideal
i: [1,642,0,61,7,1,0,0] t44_bciideal
i: [1,641,0,49,6,1,0,0] t44_bciideal
i: [1,640,0,49,6,1,0,0] t44_bciideal
i: [1,639,0,61,7,1,0,0] t44_bciideal
i: [1,638,0,61,7,1,0,0] t44_bciideal
i: [1,637,0,49,6,1,0,0] t44_bciideal
i: [1,636,0,49,6,1,0,0] t44_bciideal
i: [1,635,0,45,6,1,0,0] t44_bciideal
i: [1,634,0,45,6,1,0,0] t44_bciideal
i: [1,633,0,33,5,1,0,0] t44_bciideal
i: [1,632,0,33,5,1,0,0] t44_bciideal
i: [1,631,0,45,6,1,0,0] t44_bciideal
i: [1,630,0,45,6,1,0,0] t44_bciideal
i: [1,629,0,33,5,1,0,0] t44_bciideal
i: [1,628,0,33,5,1,0,0] t44_bciideal
i: [1,627,0,43,6,1,0,0] t44_bciideal
i: [1,626,0,43,6,1,0,0] t44_bciideal
i: [1,625,0,31,5,1,0,0] t44_bciideal
i: [1,624,0,31,5,1,0,0] t44_bciideal
i: [1,623,0,43,6,1,0,0] t44_bciideal
i: [1,622,0,43,6,1,0,0] t44_bciideal
i: [1,621,0,31,5,1,0,0] t44_bciideal
i: [1,620,0,31,5,1,0,0] t44_bciideal
i: [1,619,0,34,6,1,0,0] t44_bciideal
i: [1,618,0,34,6,1,0,0] t44_bciideal
i: [1,617,0,22,5,1,0,0] t44_bciideal
i: [1,616,0,22,5,1,0,0] t44_bciideal
i: [1,615,0,34,6,1,0,0] t44_bciideal
i: [1,614,0,34,6,1,0,0] t44_bciideal
i: [1,613,0,22,5,1,0,0] t44_bciideal
i: [1,612,0,22,5,1,0,0] t44_bciideal
i: [1,611,0,32,6,1,0,0] t44_bciideal
i: [1,610,0,32,6,1,0,0] t44_bciideal
i: [1,609,0,20,5,1,0,0] t44_bciideal
i: [1,608,0,20,5,1,0,0] t44_bciideal
i: [1,607,0,32,6,1,0,0] t44_bciideal
i: [1,606,0,32,6,1,0,0] t44_bciideal
i: [1,605,0,20,5,1,0,0] t44_bciideal
i: [1,604,0,20,5,1,0,0] t44_bciideal
i: [1,603,0,34,6,1,0,0] t44_bciideal
i: [1,602,0,34,6,1,0,0] t44_bciideal
i: [1,601,0,22,5,1,0,0] t44_bciideal
i: [1,600,0,22,5,1,0,0] t44_bciideal
i: [1,599,0,34,6,1,0,0] t44_bciideal
i: [1,598,0,34,6,1,0,0] t44_bciideal
i: [1,597,0,22,5,1,0,0] t44_bciideal
i: [1,596,0,22,5,1,0,0] t44_bciideal
i: [1,595,0,32,6,1,0,0] t44_bciideal
i: [1,594,0,32,6,1,0,0] t44_bciideal
i: [1,593,0,20,5,1,0,0] t44_bciideal
i: [1,592,0,20,5,1,0,0] t44_bciideal
i: [1,591,0,32,6,1,0,0] t44_bciideal
i: [1,590,0,32,6,1,0,0] t44_bciideal
i: [1,589,0,20,5,1,0,0] t44_bciideal
i: [1,588,0,20,5,1,0,0] t44_bciideal
i: [1,587,0,34,6,1,0,0] t44_bciideal
i: [1,586,0,34,6,1,0,0] t44_bciideal
i: [1,585,0,22,5,1,0,0] t44_bciideal
i: [1,584,0,22,5,1,0,0] t44_bciideal
i: [1,583,0,34,6,1,0,0] t44_bciideal
i: [1,582,0,34,6,1,0,0] t44_bciideal
i: [1,581,0,22,5,1,0,0] t44_bciideal
i: [1,580,0,22,5,1,0,0] t44_bciideal
i: [1,579,0,32,6,1,0,0] t44_bciideal
i: [1,578,0,32,6,1,0,0] t44_bciideal
i: [1,577,0,20,5,1,0,0] t44_bciideal
i: [1,576,0,20,5,1,0,0] t44_bciideal
i: [1,575,0,32,6,1,0,0] t44_bciideal
i: [1,574,0,32,6,1,0,0] t44_bciideal
i: [1,573,0,20,5,1,0,0] t44_bciideal
i: [1,572,0,20,5,1,0,0] t44_bciideal
i: [1,571,0,34,6,1,0,0] t44_bciideal
i: [1,570,0,34,6,1,0,0] t44_bciideal
i: [1,569,0,22,5,1,0,0] t44_bciideal
i: [1,568,0,22,5,1,0,0] t44_bciideal
i: [1,567,0,34,6,1,0,0] t44_bciideal
i: [1,566,0,34,6,1,0,0] t44_bciideal
i: [1,565,0,22,5,1,0,0] t44_bciideal
i: [1,564,0,22,5,1,0,0] t44_bciideal
i: [1,563,0,32,6,1,0,0] t44_bciideal
i: [1,562,0,32,6,1,0,0] t44_bciideal
i: [1,561,0,20,5,1,0,0] t44_bciideal
i: [1,560,0,20,5,1,0,0] t44_bciideal
i: [1,559,0,32,6,1,0,0] t44_bciideal
i: [1,558,0,32,6,1,0,0] t44_bciideal
i: [1,557,0,20,5,1,0,0] t44_bciideal
i: [1,556,0,20,5,1,0,0] t44_bciideal
i: [1,555,0,45,6,1,0,0] t44_bciideal
i: [1,554,0,45,6,1,0,0] t44_bciideal
i: [1,553,0,33,5,1,0,0] t44_bciideal
i: [1,552,0,33,5,1,0,0] t44_bciideal
i: [1,551,0,45,6,1,0,0] t44_bciideal
i: [1,550,0,45,6,1,0,0] t44_bciideal
i: [1,549,0,33,5,1,0,0] t44_bciideal
i: [1,548,0,33,5,1,0,0] t44_bciideal
i: [1,547,0,43,6,1,0,0] t44_bciideal
i: [1,546,0,43,6,1,0,0] t44_bciideal
i: [1,545,0,31,5,1,0,0] t44_bciideal
i: [1,544,0,31,5,1,0,0] t44_bciideal
i: [1,543,0,43,6,1,0,0] t44_bciideal
i: [1,542,0,43,6,1,0,0] t44_bciideal
i: [1,541,0,31,5,1,0,0] t44_bciideal
i: [1,540,0,31,5,1,0,0] t44_bciideal
i: [1,539,0,34,6,1,0,0] t44_bciideal
i: [1,538,0,34,6,1,0,0] t44_bciideal
i: [1,537,0,22,5,1,0,0] t44_bciideal
i: [1,536,0,22,5,1,0,0] t44_bciideal
i: [1,535,0,34,6,1,0,0] t44_bciideal
i: [1,534,0,34,6,1,0,0] t44_bciideal
i: [1,533,0,22,5,1,0,0] t44_bciideal
i: [1,532,0,22,5,1,0,0] t44_bciideal
i: [1,531,0,32,6,1,0,0] t44_bciideal
i: [1,530,0,32,6,1,0,0] t44_bciideal
i: [1,529,0,20,5,1,0,0] t44_bciideal
i: [1,528,0,20,5,1,0,0] t44_bciideal
i: [1,527,0,32,6,1,0,0] t44_bciideal
i: [1,526,0,32,6,1,0,0] t44_bciideal
i: [1,525,0,20,5,1,0,0] t44_bciideal
i: [1,524,0,20,5,1,0,0] t44_bciideal
i: [1,523,0,34,6,1,0,0] t44_bciideal
i: [1,522,0,34,6,1,0,0] t44_bciideal
i: [1,521,0,22,5,1,0,0] t44_bciideal
i: [1,520,0,22,5,1,0,0] t44_bciideal
i: [1,519,0,34,6,1,0,0] t44_bciideal
i: [1,518,0,34,6,1,0,0] t44_bciideal
i: [1,517,0,22,5,1,0,0] t44_bciideal
i: [1,516,0,22,5,1,0,0] t44_bciideal
i: [1,515,0,32,6,1,0,0] t44_bciideal
i: [1,514,0,32,6,1,0,0] t44_bciideal
i: [1,513,0,20,5,1,0,0] t44_bciideal
i: [1,512,0,20,5,1,0,0] t44_bciideal
i: [1,511,0,32,6,1,0,0] t44_bciideal
i: [1,510,0,32,6,1,0,0] t44_bciideal
i: [1,509,0,20,5,1,0,0] t44_bciideal
i: [1,508,0,20,5,1,0,0] t44_bciideal
i: [1,507,0,34,6,1,0,0] t44_bciideal
i: [1,506,0,34,6,1,0,0] t44_bciideal
i: [1,505,0,22,5,1,0,0] t44_bciideal
i: [1,504,0,22,5,1,0,0] t44_bciideal
i: [1,503,0,34,6,1,0,0] t44_bciideal
i: [1,502,0,34,6,1,0,0] t44_bciideal
i: [1,501,0,22,5,1,0,0] t44_bciideal
i: [1,500,0,22,5,1,0,0] t44_bciideal
i: [1,499,0,32,6,1,0,0] t44_bciideal
i: [1,498,0,32,6,1,0,0] t44_bciideal
i: [1,497,0,20,5,1,0,0] t44_bciideal
i: [1,496,0,20,5,1,0,0] t44_bciideal
i: [1,495,0,32,6,1,0,0] t44_bciideal
i: [1,494,0,32,6,1,0,0] t44_bciideal
i: [1,493,0,20,5,1,0,0] t44_bciideal
i: [1,492,0,20,5,1,0,0] t44_bciideal
i: [1,491,0,34,6,1,0,0] t44_bciideal
i: [1,490,0,34,6,1,0,0] t44_bciideal
i: [1,489,0,22,5,1,0,0] t44_bciideal
i: [1,488,0,22,5,1,0,0] t44_bciideal
i: [1,487,0,34,6,1,0,0] t44_bciideal
i: [1,486,0,34,6,1,0,0] t44_bciideal
i: [1,485,0,22,5,1,0,0] t44_bciideal
i: [1,484,0,22,5,1,0,0] t44_bciideal
i: [1,483,0,32,6,1,0,0] t44_bciideal
i: [1,482,0,32,6,1,0,0] t44_bciideal
i: [1,481,0,20,5,1,0,0] t44_bciideal
i: [1,480,0,20,5,1,0,0] t44_bciideal
i: [1,479,0,32,6,1,0,0] t44_bciideal
i: [1,478,0,32,6,1,0,0] t44_bciideal
i: [1,477,0,20,5,1,0,0] t44_bciideal
i: [1,476,0,20,5,1,0,0] t44_bciideal
i: [1,474,0,74,7,1,0,0] t44_bciideal
i: [1,473,0,62,6,1,0,0] t44_bciideal
i: [1,472,0,62,6,1,0,0] t44_bciideal
i: [1,471,0,74,7,1,0,0] t44_bciideal
i: [1,470,0,74,7,1,0,0] t44_bciideal
i: [1,469,0,62,6,1,0,0] t44_bciideal
i: [1,468,0,62,6,1,0,0] t44_bciideal
i: [1,467,0,72,7,1,0,0] t44_bciideal
i: [1,466,0,72,7,1,0,0] t44_bciideal
i: [1,465,0,60,6,1,0,0] t44_bciideal
i: [1,464,0,60,6,1,0,0] t44_bciideal
i: [1,463,0,72,7,1,0,0] t44_bciideal
i: [1,462,0,72,7,1,0,0] t44_bciideal
i: [1,461,0,60,6,1,0,0] t44_bciideal
i: [1,460,0,60,6,1,0,0] t44_bciideal
i: [1,459,0,63,7,1,0,0] t44_bciideal
i: [1,458,0,63,7,1,0,0] t44_bciideal
i: [1,457,0,51,6,1,0,0] t44_bciideal
i: [1,456,0,51,6,1,0,0] t44_bciideal
i: [1,455,0,63,7,1,0,0] t44_bciideal
i: [1,454,0,63,7,1,0,0] t44_bciideal
i: [1,453,0,51,6,1,0,0] t44_bciideal
i: [1,452,0,51,6,1,0,0] t44_bciideal
i: [1,451,0,61,7,1,0,0] t44_bciideal
i: [1,450,0,61,7,1,0,0] t44_bciideal
i: [1,449,0,49,6,1,0,0] t44_bciideal
i: [1,448,0,49,6,1,0,0] t44_bciideal
i: [1,447,0,61,7,1,0,0] t44_bciideal
i: [1,446,0,61,7,1,0,0] t44_bciideal
i: [1,445,0,49,6,1,0,0] t44_bciideal
i: [1,444,0,49,6,1,0,0] t44_bciideal
i: [1,443,0,63,7,1,0,0] t44_bciideal
i: [1,442,0,63,7,1,0,0] t44_bciideal
i: [1,441,0,51,6,1,0,0] t44_bciideal
i: [1,440,0,51,6,1,0,0] t44_bciideal
i: [1,439,0,63,7,1,0,0] t44_bciideal
i: [1,438,0,63,7,1,0,0] t44_bciideal
i: [1,437,0,51,6,1,0,0] t44_bciideal
i: [1,436,0,51,6,1,0,0] t44_bciideal
i: [1,435,0,61,7,1,0,0] t44_bciideal
i: [1,434,0,61,7,1,0,0] t44_bciideal
i: [1,433,0,49,6,1,0,0] t44_bciideal
i: [1,432,0,49,6,1,0,0] t44_bciideal
i: [1,431,0,61,7,1,0,0] t44_bciideal
i: [1,430,0,61,7,1,0,0] t44_bciideal
i: [1,429,0,49,6,1,0,0] t44_bciideal
i: [1,428,0,49,6,1,0,0] t44_bciideal
i: [1,427,0,63,7,1,0,0] t44_bciideal
i: [1,426,0,63,7,1,0,0] t44_bciideal
i: [1,425,0,51,6,1,0,0] t44_bciideal
i: [1,424,0,51,6,1,0,0] t44_bciideal
i: [1,423,0,63,7,1,0,0] t44_bciideal
i: [1,422,0,63,7,1,0,0] t44_bciideal
i: [1,421,0,51,6,1,0,0] t44_bciideal
i: [1,420,0,51,6,1,0,0] t44_bciideal
i: [1,419,0,61,7,1,0,0] t44_bciideal
i: [1,418,0,61,7,1,0,0] t44_bciideal
i: [1,417,0,49,6,1,0,0] t44_bciideal
i: [1,416,0,49,6,1,0,0] t44_bciideal
i: [1,415,0,61,7,1,0,0] t44_bciideal
i: [1,414,0,61,7,1,0,0] t44_bciideal
i: [1,413,0,49,6,1,0,0] t44_bciideal
i: [1,412,0,49,6,1,0,0] t44_bciideal
i: [1,411,0,63,7,1,0,0] t44_bciideal
i: [1,410,0,63,7,1,0,0] t44_bciideal
i: [1,409,0,51,6,1,0,0] t44_bciideal
i: [1,408,0,51,6,1,0,0] t44_bciideal
i: [1,407,0,63,7,1,0,0] t44_bciideal
i: [1,406,0,63,7,1,0,0] t44_bciideal
i: [1,405,0,51,6,1,0,0] t44_bciideal
i: [1,404,0,51,6,1,0,0] t44_bciideal
i: [1,403,0,61,7,1,0,0] t44_bciideal
i: [1,402,0,61,7,1,0,0] t44_bciideal
i: [1,401,0,49,6,1,0,0] t44_bciideal
i: [1,400,0,49,6,1,0,0] t44_bciideal
i: [1,399,0,61,7,1,0,0] t44_bciideal
i: [1,398,0,61,7,1,0,0] t44_bciideal
i: [1,397,0,49,6,1,0,0] t44_bciideal
i: [1,396,0,49,6,1,0,0] t44_bciideal
i: [1,395,0,45,6,1,0,0] t44_bciideal
i: [1,394,0,45,6,1,0,0] t44_bciideal
i: [1,393,0,33,5,1,0,0] t44_bciideal
i: [1,392,0,33,5,1,0,0] t44_bciideal
i: [1,391,0,45,6,1,0,0] t44_bciideal
i: [1,390,0,45,6,1,0,0] t44_bciideal
i: [1,389,0,33,5,1,0,0] t44_bciideal
i: [1,388,0,33,5,1,0,0] t44_bciideal
i: [1,387,0,43,6,1,0,0] t44_bciideal
i: [1,386,0,43,6,1,0,0] t44_bciideal
i: [1,385,0,31,5,1,0,0] t44_bciideal
i: [1,384,0,31,5,1,0,0] t44_bciideal
i: [1,383,0,43,6,1,0,0] t44_bciideal
i: [1,382,0,43,6,1,0,0] t44_bciideal
i: [1,381,0,31,5,1,0,0] t44_bciideal
i: [1,380,0,31,5,1,0,0] t44_bciideal
i: [1,379,0,34,6,1,0,0] t44_bciideal
i: [1,378,0,34,6,1,0,0] t44_bciideal
i: [1,377,0,22,5,1,0,0] t44_bciideal
i: [1,376,0,22,5,1,0,0] t44_bciideal
i: [1,375,0,34,6,1,0,0] t44_bciideal
i: [1,374,0,34,6,1,0,0] t44_bciideal
i: [1,373,0,22,5,1,0,0] t44_bciideal
i: [1,372,0,22,5,1,0,0] t44_bciideal
i: [1,371,0,32,6,1,0,0] t44_bciideal
i: [1,370,0,32,6,1,0,0] t44_bciideal
i: [1,369,0,20,5,1,0,0] t44_bciideal
i: [1,368,0,20,5,1,0,0] t44_bciideal
i: [1,367,0,32,6,1,0,0] t44_bciideal
i: [1,366,0,32,6,1,0,0] t44_bciideal
i: [1,365,0,20,5,1,0,0] t44_bciideal
i: [1,364,0,20,5,1,0,0] t44_bciideal
i: [1,363,0,34,6,1,0,0] t44_bciideal
i: [1,362,0,34,6,1,0,0] t44_bciideal
i: [1,361,0,22,5,1,0,0] t44_bciideal
i: [1,360,0,22,5,1,0,0] t44_bciideal
i: [1,359,0,34,6,1,0,0] t44_bciideal
i: [1,358,0,34,6,1,0,0] t44_bciideal
i: [1,357,0,22,5,1,0,0] t44_bciideal
i: [1,356,0,22,5,1,0,0] t44_bciideal
i: [1,355,0,32,6,1,0,0] t44_bciideal
i: [1,354,0,32,6,1,0,0] t44_bciideal
i: [1,353,0,20,5,1,0,0] t44_bciideal
i: [1,352,0,20,5,1,0,0] t44_bciideal
i: [1,351,0,32,6,1,0,0] t44_bciideal
i: [1,350,0,32,6,1,0,0] t44_bciideal
i: [1,349,0,20,5,1,0,0] t44_bciideal
i: [1,348,0,20,5,1,0,0] t44_bciideal
i: [1,347,0,34,6,1,0,0] t44_bciideal
i: [1,346,0,34,6,1,0,0] t44_bciideal
i: [1,345,0,22,5,1,0,0] t44_bciideal
i: [1,344,0,22,5,1,0,0] t44_bciideal
i: [1,343,0,34,6,1,0,0] t44_bciideal
i: [1,342,0,34,6,1,0,0] t44_bciideal
i: [1,341,0,22,5,1,0,0] t44_bciideal
i: [1,340,0,22,5,1,0,0] t44_bciideal
i: [1,339,0,32,6,1,0,0] t44_bciideal
i: [1,338,0,32,6,1,0,0] t44_bciideal
i: [1,337,0,20,5,1,0,0] t44_bciideal
i: [1,336,0,20,5,1,0,0] t44_bciideal
i: [1,335,0,32,6,1,0,0] t44_bciideal
i: [1,334,0,32,6,1,0,0] t44_bciideal
i: [1,333,0,20,5,1,0,0] t44_bciideal
i: [1,332,0,20,5,1,0,0] t44_bciideal
i: [1,331,0,34,6,1,0,0] t44_bciideal
i: [1,330,0,34,6,1,0,0] t44_bciideal
i: [1,329,0,22,5,1,0,0] t44_bciideal
i: [1,328,0,22,5,1,0,0] t44_bciideal
i: [1,327,0,34,6,1,0,0] t44_bciideal
i: [1,326,0,34,6,1,0,0] t44_bciideal
i: [1,325,0,22,5,1,0,0] t44_bciideal
i: [1,324,0,22,5,1,0,0] t44_bciideal
i: [1,323,0,32,6,1,0,0] t44_bciideal
i: [1,322,0,32,6,1,0,0] t44_bciideal
i: [1,321,0,20,5,1,0,0] t44_bciideal
i: [1,320,0,20,5,1,0,0] t44_bciideal
i: [1,319,0,32,6,1,0,0] t44_bciideal
i: [1,318,0,32,6,1,0,0] t44_bciideal
i: [1,317,0,20,5,1,0,0] t44_bciideal
i: [1,316,0,20,5,1,0,0] t44_bciideal
i: [1,315,0,45,6,1,0,0] t44_bciideal
i: [1,314,0,45,6,1,0,0] t44_bciideal
i: [1,313,0,33,5,1,0,0] t44_bciideal
i: [1,312,0,33,5,1,0,0] t44_bciideal
i: [1,311,0,45,6,1,0,0] t44_bciideal
i: [1,310,0,45,6,1,0,0] t44_bciideal
i: [1,309,0,33,5,1,0,0] t44_bciideal
i: [1,308,0,33,5,1,0,0] t44_bciideal
i: [1,307,0,43,6,1,0,0] t44_bciideal
i: [1,306,0,43,6,1,0,0] t44_bciideal
i: [1,305,0,31,5,1,0,0] t44_bciideal
i: [1,304,0,31,5,1,0,0] t44_bciideal
i: [1,303,0,43,6,1,0,0] t44_bciideal
i: [1,302,0,43,6,1,0,0] t44_bciideal
i: [1,301,0,31,5,1,0,0] t44_bciideal
i: [1,300,0,31,5,1,0,0] t44_bciideal
i: [1,299,0,34,6,1,0,0] t44_bciideal
i: [1,298,0,34,6,1,0,0] t44_bciideal
i: [1,297,0,22,5,1,0,0] t44_bciideal
i: [1,296,0,22,5,1,0,0] t44_bciideal
i: [1,295,0,34,6,1,0,0] t44_bciideal
i: [1,294,0,34,6,1,0,0] t44_bciideal
i: [1,293,0,22,5,1,0,0] t44_bciideal
i: [1,292,0,22,5,1,0,0] t44_bciideal
i: [1,291,0,32,6,1,0,0] t44_bciideal
i: [1,290,0,32,6,1,0,0] t44_bciideal
i: [1,289,0,20,5,1,0,0] t44_bciideal
i: [1,288,0,20,5,1,0,0] t44_bciideal
i: [1,287,0,32,6,1,0,0] t44_bciideal
i: [1,286,0,32,6,1,0,0] t44_bciideal
i: [1,285,0,20,5,1,0,0] t44_bciideal
i: [1,284,0,20,5,1,0,0] t44_bciideal
i: [1,283,0,34,6,1,0,0] t44_bciideal
i: [1,282,0,34,6,1,0,0] t44_bciideal
i: [1,281,0,22,5,1,0,0] t44_bciideal
i: [1,280,0,22,5,1,0,0] t44_bciideal
i: [1,279,0,34,6,1,0,0] t44_bciideal
i: [1,278,0,34,6,1,0,0] t44_bciideal
i: [1,277,0,22,5,1,0,0] t44_bciideal
i: [1,276,0,22,5,1,0,0] t44_bciideal
i: [1,275,0,32,6,1,0,0] t44_bciideal
i: [1,274,0,32,6,1,0,0] t44_bciideal
i: [1,273,0,20,5,1,0,0] t44_bciideal
i: [1,272,0,20,5,1,0,0] t44_bciideal
i: [1,271,0,32,6,1,0,0] t44_bciideal
i: [1,270,0,32,6,1,0,0] t44_bciideal
i: [1,269,0,20,5,1,0,0] t44_bciideal
i: [1,268,0,20,5,1,0,0] t44_bciideal
i: [1,267,0,34,6,1,0,0] t44_bciideal
i: [1,266,0,34,6,1,0,0] t44_bciideal
i: [1,265,0,22,5,1,0,0] t44_bciideal
i: [1,264,0,22,5,1,0,0] t44_bciideal
i: [1,263,0,34,6,1,0,0] t44_bciideal
i: [1,262,0,34,6,1,0,0] t44_bciideal
i: [1,261,0,22,5,1,0,0] t44_bciideal
i: [1,260,0,22,5,1,0,0] t44_bciideal
i: [1,259,0,32,6,1,0,0] t44_bciideal
i: [1,258,0,32,6,1,0,0] t44_bciideal
i: [1,257,0,20,5,1,0,0] t44_bciideal
i: [1,256,0,20,5,1,0,0] t44_bciideal
i: [1,255,0,32,6,1,0,0] t44_bciideal
i: [1,254,0,32,6,1,0,0] t44_bciideal
i: [1,253,0,20,5,1,0,0] t44_bciideal
i: [1,252,0,20,5,1,0,0] t44_bciideal
i: [1,251,0,34,6,1,0,0] t44_bciideal
i: [1,250,0,34,6,1,0,0] t44_bciideal
i: [1,249,0,22,5,1,0,0] t44_bciideal
i: [1,248,0,22,5,1,0,0] t44_bciideal
i: [1,247,0,34,6,1,0,0] t44_bciideal
i: [1,246,0,34,6,1,0,0] t44_bciideal
i: [1,245,0,22,5,1,0,0] t44_bciideal
i: [1,244,0,22,5,1,0,0] t44_bciideal
i: [1,243,0,32,6,1,0,0] t44_bciideal
i: [1,242,0,32,6,1,0,0] t44_bciideal
i: [1,241,0,20,5,1,0,0] t44_bciideal
i: [1,240,0,20,5,1,0,0] t44_bciideal
i: [1,239,0,32,6,1,0,0] t44_bciideal
i: [1,238,0,32,6,1,0,0] t44_bciideal
i: [1,237,0,20,5,1,0,0] t44_bciideal
i: [1,236,0,20,5,1,0,0] t44_bciideal
d: [2,3446,1,20,9,4,21,3434,1669]
d: [2,3447,1,18,8,5,21,3446,2843]
d: [2,3448,1,16,7,6,21,3447,3318]
d: [2,3449,1,14,6,7,21,3448,2848]
d: [2,3450,1,12,5,8,21,3449,2853]
d: [2,3451,1,10,4,9,21,3450,2858]
d: [2,3452,1,8,3,10,21,3451,2863]
d: [2,3453,1,6,2,11,21,3452,2868]
d: [2,3454,1,4,1,12,21,3453,2873]
d: [2,3347,1,33,11,1,44,1561,1540]
d: [2,3352,1,24,10,1,19,3347]
d: [2,3457,1,22,9,2,21,3352,2843]
d: [2,3458,1,20,8,3,21,3457,3318]
d: [2,3459,1,18,7,4,21,3458,2848]
d: [2,3460,1,16,6,5,21,3459,2853]
d: [2,3461,1,14,5,6,21,3460,2858]
d: [2,3462,1,12,4,7,21,3461,2863]
d: [2,3463,1,10,3,8,21,3462,2868]
d: [2,3464,1,8,2,9,21,3463,2873]
d: [2,3366,1,35,12,1,44,1538,1548]
d: [2,3369,1,30,11,1,19,3366]
d: [2,3659,1,24,10,2,21,3369,1545]
d: [2,3660,1,22,9,3,21,3659,2843]
d: [2,3661,1,20,8,4,21,3660,3318]
d: [2,3662,1,18,7,5,21,3661,2848]
d: [2,3663,1,16,6,6,21,3662,2853]
d: [2,3664,1,14,5,7,21,3663,2858]
d: [2,3665,1,12,4,8,21,3664,2863]
d: [2,3666,1,10,3,9,21,3665,2868]
d: [2,3667,1,8,2,10,21,3666,2873]
d: [2,3756,1,37,12,1,44,1537,1544]
d: [2,3762,1,34,11,1,19,3756]
d: [2,3791,1,30,10,2,21,3762,1575]
d: [2,3792,1,26,9,3,21,3791,1580]
d: [2,3793,1,24,8,4,21,3792,2843]
d: [2,3794,1,22,7,5,21,3793,3318]
d: [2,3795,1,20,6,6,21,3794,2848]
d: [2,3796,1,18,5,7,21,3795,2853]
d: [2,3797,1,16,4,8,21,3796,2858]
d: [2,3798,1,14,3,9,21,3797,2863]
d: [2,3799,1,12,2,10,21,3798,2868]
d: [2,3800,1,10,1,11,21,3799,2873]
d: [2,3336,1,33,12,1,44,1539,1585]
d: [2,3338,1,30,11,1,19,3336]
d: [2,3999,1,24,10,2,21,3338,1553]
d: [2,4000,1,20,9,3,21,3999,1601]
d: [2,4001,1,18,8,4,21,4000,2843]
d: [2,4002,1,16,7,5,21,4001,3318]
d: [2,4003,1,14,6,6,21,4002,2848]
d: [2,4004,1,12,5,7,21,4003,2853]
d: [2,4005,1,10,4,8,21,4004,2858]
d: [2,4006,1,8,3,9,21,4005,2863]
d: [2,4007,1,6,2,10,21,4006,2868]
d: [2,4008,1,4,1,11,21,4007,2873]
d: [2,4077,1,18,8,5,21,3901,2843]
d: [2,4078,1,16,7,6,21,4077,3318]
d: [2,4079,1,14,6,7,21,4078,2848]
d: [2,4080,1,12,5,8,21,4079,2853]
d: [2,4081,1,10,4,9,21,4080,2858]
d: [2,4082,1,8,3,10,21,4081,2863]
d: [2,4083,1,6,2,11,21,4082,2868]
d: [2,4084,1,4,1,12,21,4083,2873]
f: 4084,4008,3800,3667,3612,3464,3454,3314,227,228,229,230,231,232,233,234,235,236,237,238,239,240,241,242,243,244,245,246,247,248,249,250,251,252,253,254,255,256,257,258,259,260,261,262,263,264,265,266,267,268,269,270,271,272,273,274,275,276,277,278,279,280,281,282,283,284,285,286,287,288,289,290,291,292,293,294,295,296,297,298,299,300,301,302,303,304,305,306,307,308,309,310,311,312,313,314,315,316,317,318,319,320,321,322,323,324,325,326,327,328,329,330,331,332,333,334,335,336,337,338,339,340,341,342,343,344,345,346,347,348,349,350,351,352,353,354,355,356,357,358,359,360,361,362,363,364,365,366,367,368,369,370,371,372,373,374,375,376,377,378,379,380,381,382,383,384,385,386,387,388,389,390,391,392,393,394,395,396,397,398,399,400,401,402,403,404,405,406,407,408,409,410,411,412,413,414,415,416,417,418,419,420,421,422,423,424,425,426,427,428,429,430,431,432,433,434,435,436,437,438,439,440,441,442,443,444,445,446,447,448,449,450,451,452,453,454,455,456,457,458,459,460,461,462,463,464,465,466,467,468,469,470,471,472,473,474,475,476,477,478,479,480,481,482,483,484,485,486,487,488,489,490,491,492,493,494,495,496,497,498,499,500,501,502,503,504,505,506,507,508,509,510,511,512,513,514,515,516,517,518,519,520,521,522,523,524,525,526,527,528,529,530,531,532,533,534,535,536,537,538,539,540,541,542,543,544,545,546,547,548,549,550,551,552,553,554,555,556,557,558,559,560,561,562,563,564,565,566,567,568,569,570,571,572,573,574,575,576,577,578,579,580,581,582,583,584,585,586,587,588,589,590,591,592,593,594,595,596,597,598,599,600,601,602,603,604,605,606,607,608,609,610,611,612,613,614,615,616,617,618,619,620,621,622,623,624,625,626,627,628,629,630,631,632,633,634,635,636,637,638,639,640,641,642,643,644,645,646,647,648,649,650,651,652,653,654,655,656,657,658,659,660,661,662,663,664,665,666,667,668,669,670,671,672,673,674,675,676,677,678,679,680,681,682,683,684,685,686,687,688,689,690,691,692,693,694,695,696,697,698,699,700,701,702,703,704,705,706,707,708,709,710,711,712,713,714,715,716,717,718,719,720,721,722,723,724,725,726,727,728,729,730,731,732,733,734,735,736,737,738,739,740,741,742,743,744,745,746,747,748,749,750,751,752,753,754,755,756,757,758,759,760,761,762,763,764,765,766,767,768,769,770,771,772,773,774,775,776,777,778,779,780,781,782,783,784,785,786,787,788,789,790,791,792,793,794,795,796,797,798,799,800,801,802,803,804,805,806,807,808,809,810,811,812,813,814,815,816,817,818,819,820,821,822,823,824,825,826,827,828,829,830,831,832,833,834,835,836,837,838,839,840,841,842,843,844,845,846,847,848,849,850,851,852,853,854,855,856,857,858,859,860,861,862,863,864,865,866,867,868,869,870,871,872,873,874,875,876,877,878,879,880,881,882,883,884,885,886,887,888,889,890,891,892,893,894,895,896,897,898,899,900,901,902,903,904,905,906,907,908,909,910,911,912,913,914,915,916,917,918,919,920,921,922,923,924,925,926,927,928,929,930,931,932,933,934,935,936,937,938,939,940,941,942,943,944,945,946,947,948,949,950,951,952,953,954,955,956,957,958,959,960,961,962,963,964,965,966,967,968,969,970,971,972,973,974,975,976,977,978,979,980,981,982,983,984,985,986,987,988,989,990,991,992,993,994,995,996,997,998,999,1000,1001,1002,1003,1004,1005,1006,1007,1008,1009,1010,1011,1012,1013,1014,1015,1016,1017,1018,1019,1020,1021,1022,1023,1024,1025,1026,1027,1028,1029,1030,1031,1032,1033,1034,1035,1036,1037,1038,1039,1040,1041,1042,1043,1044,1045,1046,1047,1048,1049,1050,1051,1052,1053,1054,1055,1056,1057,1058,1059,1060,1061,1062,1063,1064,1065,1066,1067,1068,1069,1070,1071,1072,1073,1074,1075,1076,1077,1078,1079,1080,1081,1082,1083,1084,1085,1086,1087,1088,1089,1090,1091,1092,1093,1094,1095,1096,1097,1098,1099,1100,1101,1102,1103,1104,1105,1106,1107,1108,1109,1110,1111,1112,1113,1114,1115,1116,1117,1118,1119,1120,1121,1122,1123,1124,1125,1126,1127,1128,1129,1130,1131,1132,1133,1134,1135,1136,1137,1138,1139,1140,1141,1142,1143,1144,1145,1146,1147,1148,1149,1150,1151,1152,1153,1154,1155,1156,1157,1158,1159,1160,1161,1162,1163,1164,1165,1166,1167,1168,1169,1170,1171,1172,1173,1174,1175,1176,1177,1178,1179,1180,1181,1182,1183,1184,1185,1186,1187,1188,1189,1190,1191,1192,1193,1194,1195,1196,1197,1198,1199,1200,1201,1202,1203,1204,1205,1206,1207,1208,1209,1210,1211,1212,1213,1214,1215,1216,1217,1218,1219,1220,1221,1222,1223,1224,1225,1226,1227,1228,1229,1230,1231,1232,1233,1234,1235,1236,1237,1238,1239,1240,1241,1242,1243,1244,1245,1246,1247,1248,1249,1250,1251,1252,1253,1254,1255,1256,1257,1258,1259,1260,1261,1262,1263,1264,1265,1266,1267,1268,1269,1270,1271,1272,1273,1274,1275,1276,1277,1278,1279,1280,1281,1282,1283,1284,1285,1286,1287,1288,1289,1290,1291,1292,1293,1294,1295,1296,1297,1298,1299,1300,1301,1302,1303,1304,1305,1306,1307,1308,1309,1310,1311,1312,1313,1314,1315,1316,1317,1318,1319,1320,1321,1322,1323,1324,1325,1326,1327,1328,1329,1330,1331,1332,1333,1334,1335,1336,1337,1338,1339,1340,1341,1342,1343,1344,1345,1346,1347,1348,1349,1350,1351,1352,1353,1354,1355,1356,1357,1358,1359,1360,1361,1362,1363,1364,1365,1366,1367,1368,1369,1370,1371,1372,1373,1374,1375,1376,1377,1378,1379,1380,1381,1382,1383,1384,1385,1386,1387,1388,1389,1390,1391,1392,1393,1394,1395,1396,1397,1398,1399,1400,1401,1402,1403,1404,1405,1406,1407,1408,1409,1410,1411,1412,1413,1414,1415,1416,1417,1418,1419,1420,1421,1422,1423,1424,1425,1426,1427,1428,1429,1430,1431,1432,1433,1434,1435
% Refutation found. Thanks to Tanya!
% SZS status Theorem for bciideal__t44_bciideal
% SZS output start Proof for bciideal__t44_bciideal
1. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0)) => (! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X1,X3) => k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X2,X3) => k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 => k1_bcialg_1(X0,X2,X1) = X2))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 <=> k4_struct_0(X0) = k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))))))) [input t44_bciideal]
2. ~! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0)) => (! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X1,X3) => k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X2,X3) => k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 => k1_bcialg_1(X0,X2,X1) = X2))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 <=> k4_struct_0(X0) = k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))))))) [negated conjecture 1]
61. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))))))) [input t10_bcialg_3]
65. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0)) <=> (l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)))) [input t36_bciideal]
70. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 => k1_bcialg_1(X0,X2,X1) = X2))))) [input t7_bcialg_3]
72. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X2,X3) => k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)))))))) [input t8_bcialg_3]
74. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 <=> k4_struct_0(X0) = k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1))))))) [input t9_bcialg_3]
75. ~! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0)) => (! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X1,X3) => k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)))))) & ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => ! [X5] : (m1_subset_1(X5,u1_struct_0(X0)) => (k1_bcialg_1(X0,X4,X5) = k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),X4)) & k1_bcialg_1(X0,X4,X5) = k1_bcialg_1(X0,X4,k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X5,X4)))))) & ! [X6] : (m1_subset_1(X6,u1_struct_0(X0)) => ! [X7] : (m1_subset_1(X7,u1_struct_0(X0)) => ! [X8] : (m1_subset_1(X8,u1_struct_0(X0)) => (r1_bcialg_1(X0,X7,X8) => k1_bcialg_1(X0,X7,X6) = k1_bcialg_1(X0,k1_bcialg_1(X0,X8,X6),k1_bcialg_1(X0,X8,X7)))))) & ! [X9] : (m1_subset_1(X9,u1_struct_0(X0)) => ! [X10] : (m1_subset_1(X10,u1_struct_0(X0)) => (k1_bcialg_1(X0,X9,X10) = X9 => k1_bcialg_1(X0,X10,X9) = X10))) & ! [X11] : (m1_subset_1(X11,u1_struct_0(X0)) => ! [X12] : (m1_subset_1(X12,u1_struct_0(X0)) => (k1_bcialg_1(X0,X11,X12) = X11 <=> k4_struct_0(X0) = k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)))))))) [rectify 2]
89. ? [X0] : (((? [X1] : (? [X2] : (? [X3] : ((k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)) & r1_bcialg_1(X0,X1,X3)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | ? [X4] : (? [X5] : ((k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),X4)) | k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,X4,k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(X0))) & m1_subset_1(X4,u1_struct_0(X0))) | ? [X6] : (? [X7] : (? [X8] : ((k1_bcialg_1(X0,X7,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X8,X6),k1_bcialg_1(X0,X8,X7)) & r1_bcialg_1(X0,X7,X8)) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) & m1_subset_1(X6,u1_struct_0(X0))) | ? [X9] : (? [X10] : ((k1_bcialg_1(X0,X10,X9) != X10 & k1_bcialg_1(X0,X9,X10) = X9) & m1_subset_1(X10,u1_struct_0(X0))) & m1_subset_1(X9,u1_struct_0(X0))) | ? [X11] : (? [X12] : ((k1_bcialg_1(X0,X11,X12) = X11 <~> k4_struct_0(X0) = k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11))) & m1_subset_1(X12,u1_struct_0(X0))) & m1_subset_1(X11,u1_struct_0(X0)))) & (m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0))) & (l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0))) [ennf transformation 75]
90. ? [X0] : ((? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)) & r1_bcialg_1(X0,X1,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | ? [X4] : (? [X5] : ((k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),X4)) | k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,X4,k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(X0))) & m1_subset_1(X4,u1_struct_0(X0))) | ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(X0,X7,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X8,X6),k1_bcialg_1(X0,X8,X7)) & r1_bcialg_1(X0,X7,X8) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) & m1_subset_1(X6,u1_struct_0(X0))) | ? [X9] : (? [X10] : (k1_bcialg_1(X0,X10,X9) != X10 & k1_bcialg_1(X0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(X0))) & m1_subset_1(X9,u1_struct_0(X0))) | ? [X11] : (? [X12] : ((k1_bcialg_1(X0,X11,X12) = X11 <~> k4_struct_0(X0) = k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11))) & m1_subset_1(X12,u1_struct_0(X0))) & m1_subset_1(X11,u1_struct_0(X0)))) & m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0) & l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) [flattening 89]
130. ! [X0] : ((! [X1] : (! [X2] : ((k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 61]
131. ! [X0] : (! [X1] : (! [X2] : ((k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 130]
132. ! [X0] : ((! [X1] : (! [X2] : ((k1_bcialg_1(X0,X2,X1) = X2 | k1_bcialg_1(X0,X1,X2) != X1) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 70]
133. ! [X0] : (! [X1] : (! [X2] : (k1_bcialg_1(X0,X2,X1) = X2 | k1_bcialg_1(X0,X1,X2) != X1 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 132]
134. ! [X0] : ((! [X1] : (! [X2] : ((k1_bcialg_1(X0,X1,X2) = X1 <=> k4_struct_0(X0) = k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 74]
135. ! [X0] : (! [X1] : (! [X2] : ((k1_bcialg_1(X0,X1,X2) = X1 <=> k4_struct_0(X0) = k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 134]
136. ! [X0] : ((! [X1] : (! [X2] : (! [X3] : ((k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)) | ~r1_bcialg_1(X0,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 72]
137. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)) | ~r1_bcialg_1(X0,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 136]
140. ! [X0] : (((m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0)) <=> (l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 65]
141. ! [X0] : (((m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0)) <=> (l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 140]
178. ? [X0] : ((? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)) & r1_bcialg_1(X0,X1,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | ? [X4] : (? [X5] : ((k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),X4)) | k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,X4,k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(X0))) & m1_subset_1(X4,u1_struct_0(X0))) | ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(X0,X7,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X8,X6),k1_bcialg_1(X0,X8,X7)) & r1_bcialg_1(X0,X7,X8) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) & m1_subset_1(X6,u1_struct_0(X0))) | ? [X9] : (? [X10] : (k1_bcialg_1(X0,X10,X9) != X10 & k1_bcialg_1(X0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(X0))) & m1_subset_1(X9,u1_struct_0(X0))) | ? [X11] : (? [X12] : (((k4_struct_0(X0) != k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) != X11) & (k4_struct_0(X0) = k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) = X11)) & m1_subset_1(X12,u1_struct_0(X0))) & m1_subset_1(X11,u1_struct_0(X0)))) & m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0) & l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) [nnf transformation 90]
179. ? [X0] : ((? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)) & r1_bcialg_1(X0,X1,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | ? [X4] : (? [X5] : ((k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),X4)) | k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,X4,k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(X0))) & m1_subset_1(X4,u1_struct_0(X0))) | ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(X0,X7,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X8,X6),k1_bcialg_1(X0,X8,X7)) & r1_bcialg_1(X0,X7,X8) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) & m1_subset_1(X6,u1_struct_0(X0))) | ? [X9] : (? [X10] : (k1_bcialg_1(X0,X10,X9) != X10 & k1_bcialg_1(X0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(X0))) & m1_subset_1(X9,u1_struct_0(X0))) | ? [X11] : (? [X12] : ((k4_struct_0(X0) != k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) != X11) & (k4_struct_0(X0) = k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) = X11) & m1_subset_1(X12,u1_struct_0(X0))) & m1_subset_1(X11,u1_struct_0(X0)))) & m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0) & l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) [flattening 178]
180. ? [X0] : ((? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)) & r1_bcialg_1(X0,X1,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | ? [X4] : (? [X5] : ((k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),X4)) | k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,X4,k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(X0))) & m1_subset_1(X4,u1_struct_0(X0))) | ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(X0,X7,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X8,X6),k1_bcialg_1(X0,X8,X7)) & r1_bcialg_1(X0,X7,X8) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) & m1_subset_1(X6,u1_struct_0(X0))) | ? [X9] : (? [X10] : (k1_bcialg_1(X0,X10,X9) != X10 & k1_bcialg_1(X0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(X0))) & m1_subset_1(X9,u1_struct_0(X0))) | ? [X11] : (? [X12] : ((k4_struct_0(X0) != k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) != X11) & (k4_struct_0(X0) = k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) = X11) & m1_subset_1(X12,u1_struct_0(X0))) & m1_subset_1(X11,u1_struct_0(X0)))) & m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0) & l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X3,X1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X1,X2)) & r1_bcialg_1(sK0,X1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) | ? [X4] : (? [X5] : ((k1_bcialg_1(sK0,X4,X5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X4,X5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X4,X5),X4)) | k1_bcialg_1(sK0,X4,X5) != k1_bcialg_1(sK0,X4,k1_bcialg_1(sK0,X5,k1_bcialg_1(sK0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(sK0))) & m1_subset_1(X4,u1_struct_0(sK0))) | ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(sK0,X7,X6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,X6),k1_bcialg_1(sK0,X8,X7)) & r1_bcialg_1(sK0,X7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(X7,u1_struct_0(sK0))) & m1_subset_1(X6,u1_struct_0(sK0))) | ? [X9] : (? [X10] : (k1_bcialg_1(sK0,X10,X9) != X10 & k1_bcialg_1(sK0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(sK0))) & m1_subset_1(X9,u1_struct_0(sK0))) | ? [X11] : (? [X12] : ((k4_struct_0(sK0) != k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,X11)) | k1_bcialg_1(sK0,X11,X12) != X11) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,X11)) | k1_bcialg_1(sK0,X11,X12) = X11) & m1_subset_1(X12,u1_struct_0(sK0))) & m1_subset_1(X11,u1_struct_0(sK0)))) & m2_bcialg_1(k1_tarski(k4_struct_0(sK0)),sK0) & v3_bciideal(k1_tarski(k4_struct_0(sK0)),sK0) & l2_bcialg_1(sK0) & v8_bcialg_1(sK0) & v7_bcialg_1(sK0) & v5_bcialg_1(sK0) & v4_bcialg_1(sK0) & v3_bcialg_1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
181. ? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X3,X1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X1,X2)) & r1_bcialg_1(sK0,X1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,sK1,X2)) & r1_bcialg_1(sK0,sK1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
182. ? [X2] : (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,sK1,X2)) & r1_bcialg_1(sK0,sK1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,X3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,sK1,sK2)) & r1_bcialg_1(sK0,sK1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
183. ? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,X3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,sK1,sK2)) & r1_bcialg_1(sK0,sK1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) => (k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) & r1_bcialg_1(sK0,sK1,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
184. ? [X4] : (? [X5] : ((k1_bcialg_1(sK0,X4,X5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X4,X5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X4,X5),X4)) | k1_bcialg_1(sK0,X4,X5) != k1_bcialg_1(sK0,X4,k1_bcialg_1(sK0,X5,k1_bcialg_1(sK0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(sK0))) & m1_subset_1(X4,u1_struct_0(sK0))) => (? [X5] : ((k1_bcialg_1(sK0,sK4,X5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,X5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,X5),sK4)) | k1_bcialg_1(sK0,sK4,X5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,X5,k1_bcialg_1(sK0,X5,sK4)))) & m1_subset_1(X5,u1_struct_0(sK0))) & m1_subset_1(sK4,u1_struct_0(sK0))) [choice axiom]
185. ? [X5] : ((k1_bcialg_1(sK0,sK4,X5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,X5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,X5),sK4)) | k1_bcialg_1(sK0,sK4,X5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,X5,k1_bcialg_1(sK0,X5,sK4)))) & m1_subset_1(X5,u1_struct_0(sK0))) => ((k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4)))) & m1_subset_1(sK5,u1_struct_0(sK0))) [choice axiom]
186. ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(sK0,X7,X6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,X6),k1_bcialg_1(sK0,X8,X7)) & r1_bcialg_1(sK0,X7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(X7,u1_struct_0(sK0))) & m1_subset_1(X6,u1_struct_0(sK0))) => (? [X7] : (? [X8] : (k1_bcialg_1(sK0,X7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,sK6),k1_bcialg_1(sK0,X8,X7)) & r1_bcialg_1(sK0,X7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(X7,u1_struct_0(sK0))) & m1_subset_1(sK6,u1_struct_0(sK0))) [choice axiom]
187. ? [X7] : (? [X8] : (k1_bcialg_1(sK0,X7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,sK6),k1_bcialg_1(sK0,X8,X7)) & r1_bcialg_1(sK0,X7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(X7,u1_struct_0(sK0))) => (? [X8] : (k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,sK6),k1_bcialg_1(sK0,X8,sK7)) & r1_bcialg_1(sK0,sK7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(sK7,u1_struct_0(sK0))) [choice axiom]
188. ? [X8] : (k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,sK6),k1_bcialg_1(sK0,X8,sK7)) & r1_bcialg_1(sK0,sK7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) => (k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) & r1_bcialg_1(sK0,sK7,sK8) & m1_subset_1(sK8,u1_struct_0(sK0))) [choice axiom]
189. ? [X9] : (? [X10] : (k1_bcialg_1(sK0,X10,X9) != X10 & k1_bcialg_1(sK0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(sK0))) & m1_subset_1(X9,u1_struct_0(sK0))) => (? [X10] : (k1_bcialg_1(sK0,X10,sK9) != X10 & sK9 = k1_bcialg_1(sK0,sK9,X10) & m1_subset_1(X10,u1_struct_0(sK0))) & m1_subset_1(sK9,u1_struct_0(sK0))) [choice axiom]
190. ? [X10] : (k1_bcialg_1(sK0,X10,sK9) != X10 & sK9 = k1_bcialg_1(sK0,sK9,X10) & m1_subset_1(X10,u1_struct_0(sK0))) => (sK10 != k1_bcialg_1(sK0,sK10,sK9) & sK9 = k1_bcialg_1(sK0,sK9,sK10) & m1_subset_1(sK10,u1_struct_0(sK0))) [choice axiom]
191. ? [X11] : (? [X12] : ((k4_struct_0(sK0) != k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,X11)) | k1_bcialg_1(sK0,X11,X12) != X11) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,X11)) | k1_bcialg_1(sK0,X11,X12) = X11) & m1_subset_1(X12,u1_struct_0(sK0))) & m1_subset_1(X11,u1_struct_0(sK0))) => (? [X12] : ((k4_struct_0(sK0) != k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,X12)) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,X12)) & m1_subset_1(X12,u1_struct_0(sK0))) & m1_subset_1(sK11,u1_struct_0(sK0))) [choice axiom]
192. ? [X12] : ((k4_struct_0(sK0) != k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,X12)) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,X12)) & m1_subset_1(X12,u1_struct_0(sK0))) => ((k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12)) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12)) & m1_subset_1(sK12,u1_struct_0(sK0))) [choice axiom]
193. ((((k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) & r1_bcialg_1(sK0,sK1,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) | (((k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4)))) & m1_subset_1(sK5,u1_struct_0(sK0))) & m1_subset_1(sK4,u1_struct_0(sK0))) | (((k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) & r1_bcialg_1(sK0,sK7,sK8) & m1_subset_1(sK8,u1_struct_0(sK0))) & m1_subset_1(sK7,u1_struct_0(sK0))) & m1_subset_1(sK6,u1_struct_0(sK0))) | ((sK10 != k1_bcialg_1(sK0,sK10,sK9) & sK9 = k1_bcialg_1(sK0,sK9,sK10) & m1_subset_1(sK10,u1_struct_0(sK0))) & m1_subset_1(sK9,u1_struct_0(sK0))) | (((k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12)) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12)) & m1_subset_1(sK12,u1_struct_0(sK0))) & m1_subset_1(sK11,u1_struct_0(sK0)))) & m2_bcialg_1(k1_tarski(k4_struct_0(sK0)),sK0) & v3_bciideal(k1_tarski(k4_struct_0(sK0)),sK0) & l2_bcialg_1(sK0) & v8_bcialg_1(sK0) & v7_bcialg_1(sK0) & v5_bcialg_1(sK0) & v4_bcialg_1(sK0) & v3_bcialg_1(sK0) & ~v2_struct_0(sK0) [skolemisation 179,192,191,190,189,188,187,186,185,184,183,182,181,180]
200. ! [X0] : (! [X1] : (! [X2] : (((k1_bcialg_1(X0,X1,X2) = X1 | k4_struct_0(X0) != k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1))) & (k4_struct_0(X0) = k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) | k1_bcialg_1(X0,X1,X2) != X1)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [nnf transformation 135]
201. ! [X0] : ((((m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0)) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) & ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) | (~m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) | ~v3_bciideal(k1_tarski(k4_struct_0(X0)),X0)))) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [nnf transformation 141]
202. ! [X0] : ((((m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) & v3_bciideal(k1_tarski(k4_struct_0(X0)),X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) & ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) | ~m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) | ~v3_bciideal(k1_tarski(k4_struct_0(X0)),X0))) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 201]
227. ~v2_struct_0(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
228. v3_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
229. v4_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
230. v5_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
231. v7_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
232. v8_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
233. l2_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
234. v3_bciideal(k1_tarski(k4_struct_0(sK0)),sK0) [cnf transformation 193] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
235. m2_bcialg_1(k1_tarski(k4_struct_0(sK0)),sK0) [cnf transformation 193] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
236. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
237. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
238. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
239. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
240. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
241. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
242. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
243. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
244. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
245. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
246. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
247. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
248. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
249. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
250. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
251. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
252. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
253. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
254. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
255. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
256. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
257. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
258. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
259. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
260. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
261. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
262. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
263. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
264. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
265. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
266. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
267. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
268. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
269. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
270. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
271. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
272. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
273. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
274. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
275. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
276. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
277. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
278. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
279. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
280. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
281. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
282. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
283. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
284. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
285. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
286. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
287. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
288. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
289. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
290. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
291. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
292. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
293. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
294. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
295. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
296. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
297. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
298. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
299. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
300. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
301. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
302. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
303. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
304. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
305. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
306. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
307. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
308. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
309. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
310. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
311. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
312. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
313. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
314. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
315. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
316. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
317. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
318. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
319. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
320. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
321. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
322. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
323. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
324. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
325. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
326. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
327. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
328. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
329. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
330. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
331. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
332. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
333. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
334. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
335. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
336. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
337. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
338. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
339. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
340. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
341. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
342. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
343. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
344. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
345. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
346. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
347. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
348. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
349. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
350. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
351. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
352. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
353. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
354. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
355. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
356. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
357. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
358. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
359. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
360. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
361. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
362. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
363. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
364. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
365. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
366. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
367. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
368. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
369. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
370. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
371. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
372. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
373. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
374. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
375. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
376. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
377. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
378. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
379. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
380. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
381. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
382. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
383. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
384. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
385. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
386. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
387. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
388. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
389. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
390. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
391. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
392. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
393. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
394. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
395. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
396. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
397. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
398. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
399. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
400. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
401. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
402. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
403. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
404. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
405. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
406. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
407. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
408. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
409. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
410. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
411. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
412. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
413. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
414. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
415. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
416. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
417. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
418. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
419. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
420. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
421. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
422. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
423. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
424. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
425. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
426. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
427. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
428. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
429. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
430. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
431. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
432. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
433. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
434. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
435. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
436. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
437. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
438. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
439. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
440. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
441. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
442. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
443. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
444. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
445. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
446. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
447. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
448. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
449. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
450. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
451. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
452. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
453. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
454. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
455. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
456. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
457. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
458. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
459. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
460. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
461. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
462. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
463. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
464. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
465. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
466. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
467. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
468. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
469. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
470. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
471. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
472. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
473. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
474. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
475. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
476. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
477. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
478. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
479. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
480. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
481. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
482. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
483. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
484. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
485. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
486. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
487. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
488. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
489. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
490. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
491. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
492. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
493. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
494. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
495. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
496. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
497. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
498. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
499. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
500. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
501. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
502. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
503. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
504. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
505. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
506. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
507. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
508. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
509. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
510. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
511. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
512. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
513. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
514. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
515. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
516. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
517. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
518. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
519. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
520. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
521. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
522. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
523. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
524. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
525. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
526. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
527. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
528. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
529. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
530. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
531. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
532. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
533. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
534. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
535. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
536. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
537. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
538. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
539. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
540. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
541. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
542. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
543. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
544. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
545. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
546. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
547. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
548. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
549. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
550. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
551. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
552. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
553. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
554. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
555. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
556. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
557. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
558. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
559. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
560. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
561. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
562. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
563. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
564. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
565. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
566. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
567. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
568. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
569. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
570. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
571. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
572. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
573. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
574. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
575. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
576. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
577. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
578. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
579. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
580. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
581. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
582. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
583. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
584. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
585. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
586. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
587. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
588. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
589. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
590. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
591. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
592. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
593. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
594. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
595. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
596. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
597. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
598. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
599. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
600. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
601. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
602. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
603. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
604. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
605. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
606. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
607. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
608. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
609. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
610. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
611. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
612. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
613. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
614. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
615. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
616. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
617. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
618. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
619. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
620. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
621. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
622. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
623. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
624. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
625. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
626. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
627. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
628. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
629. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
630. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
631. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
632. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
633. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
634. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
635. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
636. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
637. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
638. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
639. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
640. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
641. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
642. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
643. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
644. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
645. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
646. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
647. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
648. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
649. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
650. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
651. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
652. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
653. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
654. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
655. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
656. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
657. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
658. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
659. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
660. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
661. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
662. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
663. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
664. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
665. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
666. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
667. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
668. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
669. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
670. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
671. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
672. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
673. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
674. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
675. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
676. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
677. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
678. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
679. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
680. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
681. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
682. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
683. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
684. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
685. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
686. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
687. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
688. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
689. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
690. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
691. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
692. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
693. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
694. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
695. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
696. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
697. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
698. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
699. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
700. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
701. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
702. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
703. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
704. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
705. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
706. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
707. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
708. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
709. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
710. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
711. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
712. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
713. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
714. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
715. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
716. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
717. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
718. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
719. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
720. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
721. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
722. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
723. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
724. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
725. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
726. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
727. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
728. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
729. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
730. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
731. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
732. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
733. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
734. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
735. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
736. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
737. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
738. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
739. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
740. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
741. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
742. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
743. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
744. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
745. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
746. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
747. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
748. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
749. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
750. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
751. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
752. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
753. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
754. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
755. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
756. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
757. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
758. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
759. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
760. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
761. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
762. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
763. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
764. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
765. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
766. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
767. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
768. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
769. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
770. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
771. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
772. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
773. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
774. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
775. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
776. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
777. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
778. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
779. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
780. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
781. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
782. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
783. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
784. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
785. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
786. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
787. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
788. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
789. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
790. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
791. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
792. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
793. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
794. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
795. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
796. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
797. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
798. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
799. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
800. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
801. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
802. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
803. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
804. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
805. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
806. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
807. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
808. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
809. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
810. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
811. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
812. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
813. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
814. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
815. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
816. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
817. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
818. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
819. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
820. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
821. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
822. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
823. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
824. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
825. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
826. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
827. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
828. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
829. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
830. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
831. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
832. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
833. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
834. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
835. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
836. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
837. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
838. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
839. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
840. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
841. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
842. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
843. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
844. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
845. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
846. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
847. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
848. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
849. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
850. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
851. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
852. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
853. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
854. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
855. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
856. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
857. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
858. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
859. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
860. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
861. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
862. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
863. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
864. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
865. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
866. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
867. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
868. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
869. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
870. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
871. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
872. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
873. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
874. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
875. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
876. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
877. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
878. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
879. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
880. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
881. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
882. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
883. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
884. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
885. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
886. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
887. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
888. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
889. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
890. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
891. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
892. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
893. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
894. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
895. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
896. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
897. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
898. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
899. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
900. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
901. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
902. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
903. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
904. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
905. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
906. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
907. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
908. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
909. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
910. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
911. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
912. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
913. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
914. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
915. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
916. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
917. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
918. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
919. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
920. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
921. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
922. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
923. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
924. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
925. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
926. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
927. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
928. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
929. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
930. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
931. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
932. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
933. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
934. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
935. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
936. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
937. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
938. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
939. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
940. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
941. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
942. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
943. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
944. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
945. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
946. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
947. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
948. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
949. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
950. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
951. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
952. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
953. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
954. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
955. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
956. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
957. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
958. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
959. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
960. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
961. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
962. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
963. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
964. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
965. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
966. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
967. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
968. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
969. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
970. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
971. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
972. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
973. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
974. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
975. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
976. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
977. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
978. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
979. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
980. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
981. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
982. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
983. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
984. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
985. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
986. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
987. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
988. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
989. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
990. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
991. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
992. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
993. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
994. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
995. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
996. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
997. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
998. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
999. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1000. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1001. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1002. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1003. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1004. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1005. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1006. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1007. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1008. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1009. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1010. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1011. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1012. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1013. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1014. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1015. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1016. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1017. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1018. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1019. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1020. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1021. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1022. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1023. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1024. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1025. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1026. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1027. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1028. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1029. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1030. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1031. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1032. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1033. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1034. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1035. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1036. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1037. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1038. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1039. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1040. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1041. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1042. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1043. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1044. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1045. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1046. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1047. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1048. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1049. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1050. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1051. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1052. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1053. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1054. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1055. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1056. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1057. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1058. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1059. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1060. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1061. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1062. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1063. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1064. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1065. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1066. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1067. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1068. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1069. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1070. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1071. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1072. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1073. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1074. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1075. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1076. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1077. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1078. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1079. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1080. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1081. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1082. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1083. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1084. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1085. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1086. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1087. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1088. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1089. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1090. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1091. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1092. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1093. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1094. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1095. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1096. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1097. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1098. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1099. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1100. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1101. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1102. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1103. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1104. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1105. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1106. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1107. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1108. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1109. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1110. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1111. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1112. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1113. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1114. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1115. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1116. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1117. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1118. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1119. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1120. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1121. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1122. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1123. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1124. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1125. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1126. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1127. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1128. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1129. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1130. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1131. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1132. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1133. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1134. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1135. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1136. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1137. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1138. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1139. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1140. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1141. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1142. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1143. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1144. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1145. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1146. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1147. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1148. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1149. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1150. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1151. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1152. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1153. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1154. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1155. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1156. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1157. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1158. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1159. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1160. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1161. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1162. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1163. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1164. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1165. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1166. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1167. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1168. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1169. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1170. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1171. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1172. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1173. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1174. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1175. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1176. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1177. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1178. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1179. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1180. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1181. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1182. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1183. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1184. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1185. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1186. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1187. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1188. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1189. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1190. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1191. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1192. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1193. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1194. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1195. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1196. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1197. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1198. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1199. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1200. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1201. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1202. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1203. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1204. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1205. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1206. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1207. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1208. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1209. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1210. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1211. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1212. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1213. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1214. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1215. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1216. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1217. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1218. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1219. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1220. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1221. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1222. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1223. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1224. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1225. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1226. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1227. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1228. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1229. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1230. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1231. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1232. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1233. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1234. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1235. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1236. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1237. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1238. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1239. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1240. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1241. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1242. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1243. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1244. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1245. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1246. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1247. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1248. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1249. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1250. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1251. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1252. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1253. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1254. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1255. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1256. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1257. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1258. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1259. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1260. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1261. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1262. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1263. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1264. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1265. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1266. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1267. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1268. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1269. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1270. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1271. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1272. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1273. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1274. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1275. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1276. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1277. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1278. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1279. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1280. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1281. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1282. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1283. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1284. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1285. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1286. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1287. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1288. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1289. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1290. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1291. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1292. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1293. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1294. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1295. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1296. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1297. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1298. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1299. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1300. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1301. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1302. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1303. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1304. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1305. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1306. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1307. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1308. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1309. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1310. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1311. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1312. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1313. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1314. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1315. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1316. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1317. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1318. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1319. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1320. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1321. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1322. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1323. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1324. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1325. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1326. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1327. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1328. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1329. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1330. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1331. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1332. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1333. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1334. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1335. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1336. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1337. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1338. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1339. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1340. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1341. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1342. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1343. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1344. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1345. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1346. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1347. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1348. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1349. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1350. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1351. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1352. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1353. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1354. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1355. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1356. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1357. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1358. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1359. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1360. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1361. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1362. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1363. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1364. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1365. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1366. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1367. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1368. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1369. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1370. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1371. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1372. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1373. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1374. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1375. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1376. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1377. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1378. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1379. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1380. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1381. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1382. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1383. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1384. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1385. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1386. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1387. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1388. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1389. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1390. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1391. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1392. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1393. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1394. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1395. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1396. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1397. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1398. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1399. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1400. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1401. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1402. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1403. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1404. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1405. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1406. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1407. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1408. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1409. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1410. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1411. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1412. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1413. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1414. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1415. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1416. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1417. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1418. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1419. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1420. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:83,wCS:8300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1421. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:83,wCS:8300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1422. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:95,wCS:9500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1423. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:95,wCS:9500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1424. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:83,wCS:8300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1425. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:83,wCS:8300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1426. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:95,wCS:9500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1427. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:95,wCS:9500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1428. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:85,wCS:8500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1429. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:85,wCS:8500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1430. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:97,wCS:9700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1431. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:97,wCS:9700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1432. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:85,wCS:8500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1433. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:85,wCS:8500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1434. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:97,wCS:9700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1435. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:97,wCS:9700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1472. k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 131] {a:0,w:53,wCS:21200,sine:2,thAx:0,allAx:1,thDist:-1}
1473. k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 131] {a:0,w:56,wCS:22400,sine:2,thAx:0,allAx:1,thDist:-1}
1474. k1_bcialg_1(X0,X2,X1) = X2 | k1_bcialg_1(X0,X1,X2) != X1 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 133] {a:0,w:50,wCS:20000,sine:2,thAx:0,allAx:1,thDist:-1}
1475. k4_struct_0(X0) = k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) | k1_bcialg_1(X0,X1,X2) != X1 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 200] {a:0,w:54,wCS:21600,sine:1,thAx:0,allAx:1,thDist:-1}
1476. k1_bcialg_1(X0,X1,X2) = X1 | k4_struct_0(X0) != k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 200] {a:0,w:54,wCS:21600,sine:1,thAx:0,allAx:1,thDist:-1}
1477. k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)) | ~r1_bcialg_1(X0,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 137] {a:0,w:61,wCS:24400,sine:1,thAx:0,allAx:1,thDist:-1}
1485. ~m2_bcialg_1(k1_tarski(k4_struct_0(X0)),X0) | v1_bcialg_3(X0) | ~v3_bciideal(k1_tarski(k4_struct_0(X0)),X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 202] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1536. k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)) | ~r1_bcialg_1(X0,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1477] {a:0,w:47,wCS:18800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1537. k1_bcialg_1(X0,X1,X2) != X1 | k4_struct_0(X0) = k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1475] {a:0,w:40,wCS:16000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1538. k4_struct_0(X0) != k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) | k1_bcialg_1(X0,X1,X2) = X1 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1476] {a:0,w:40,wCS:16000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1539. k1_bcialg_1(X0,X1,X2) != X1 | k1_bcialg_1(X0,X2,X1) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1474] {a:0,w:36,wCS:14400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1540. k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1472] {a:0,w:39,wCS:15600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1541. k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1473] {a:0,w:42,wCS:16800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1543. 1 <=> sK11 = k1_bcialg_1(sK0,sK11,sK12) [avatar definition]
1544. sK11 = k1_bcialg_1(sK0,sK11,sK12) <- (1) [avatar component clause 1543] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1545. sK11 != k1_bcialg_1(sK0,sK11,sK12) <- (~1) [avatar component clause 1543] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1547. 2 <=> k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) [avatar definition]
1548. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) <- (2) [avatar component clause 1547] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1551. 3 <=> sK10 = k1_bcialg_1(sK0,sK10,sK9) [avatar definition]
1553. sK10 != k1_bcialg_1(sK0,sK10,sK9) <- (~3) [avatar component clause 1551] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1555. 4 <=> k1_bcialg_1(sK0,sK7,sK6) = k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) [avatar definition]
1557. k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) <- (~4) [avatar component clause 1555] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1559. 5 <=> k1_bcialg_1(sK0,sK4,sK5) = k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) [avatar definition]
1561. k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) <- (~5) [avatar component clause 1559] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1563. 6 <=> k1_bcialg_1(sK0,sK4,sK5) = k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) [avatar definition]
1565. k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) <- (~6) [avatar component clause 1563] {a:0,w:18,wCS:1800,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1567. 7 <=> k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) = k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) [avatar definition]
1569. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) <- (~7) [avatar component clause 1567] {a:0,w:27,wCS:2700,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1570. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1435,1567,1563,1559,1555,1551,1547,1543]
1571. 1 | 2 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1434,1567,1563,1559,1555,1551,1547,1543]
1573. 8 <=> m1_subset_1(sK12,u1_struct_0(sK0)) [avatar definition]
1575. m1_subset_1(sK12,u1_struct_0(sK0)) <- (8) [avatar component clause 1573] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1576. 8 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1433,1567,1563,1559,1555,1551,1573]
1578. 9 <=> m1_subset_1(sK11,u1_struct_0(sK0)) [avatar definition]
1580. m1_subset_1(sK11,u1_struct_0(sK0)) <- (9) [avatar component clause 1578] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1581. 9 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1432,1567,1563,1559,1555,1551,1578]
1583. 10 <=> sK9 = k1_bcialg_1(sK0,sK9,sK10) [avatar definition]
1585. sK9 = k1_bcialg_1(sK0,sK9,sK10) <- (10) [avatar component clause 1583] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1586. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1431,1567,1563,1559,1555,1583,1547,1543]
1587. 1 | 2 | 10 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1430,1567,1563,1559,1555,1583,1547,1543]
1588. 8 | 10 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1429,1567,1563,1559,1555,1583,1573]
1589. 9 | 10 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1428,1567,1563,1559,1555,1583,1578]
1591. 11 <=> m1_subset_1(sK10,u1_struct_0(sK0)) [avatar definition]
1594. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1427,1567,1563,1559,1555,1591,1547,1543]
1595. 1 | 2 | 11 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1426,1567,1563,1559,1555,1591,1547,1543]
1596. 8 | 11 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1425,1567,1563,1559,1555,1591,1573]
1597. 9 | 11 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1424,1567,1563,1559,1555,1591,1578]
1599. 12 <=> m1_subset_1(sK9,u1_struct_0(sK0)) [avatar definition]
1601. m1_subset_1(sK9,u1_struct_0(sK0)) <- (12) [avatar component clause 1599] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1602. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1423,1567,1563,1559,1555,1599,1547,1543]
1603. 1 | 2 | 12 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1422,1567,1563,1559,1555,1599,1547,1543]
1604. 8 | 12 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1421,1567,1563,1559,1555,1599,1573]
1605. 9 | 12 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1420,1567,1563,1559,1555,1599,1578]
1607. 13 <=> r1_bcialg_1(sK0,sK7,sK8) [avatar definition]
1609. r1_bcialg_1(sK0,sK7,sK8) <- (13) [avatar component clause 1607] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1610. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | ~7 [avatar split clause 1419,1567,1563,1559,1607,1551,1547,1543]
1611. 1 | 2 | ~3 | 13 | ~5 | ~6 | ~7 [avatar split clause 1418,1567,1563,1559,1607,1551,1547,1543]
1612. 8 | ~3 | 13 | ~5 | ~6 | ~7 [avatar split clause 1417,1567,1563,1559,1607,1551,1573]
1613. 9 | ~3 | 13 | ~5 | ~6 | ~7 [avatar split clause 1416,1567,1563,1559,1607,1551,1578]
1614. ~1 | ~2 | 10 | 13 | ~5 | ~6 | ~7 [avatar split clause 1415,1567,1563,1559,1607,1583,1547,1543]
1615. 1 | 2 | 10 | 13 | ~5 | ~6 | ~7 [avatar split clause 1414,1567,1563,1559,1607,1583,1547,1543]
1616. 8 | 10 | 13 | ~5 | ~6 | ~7 [avatar split clause 1413,1567,1563,1559,1607,1583,1573]
1617. 9 | 10 | 13 | ~5 | ~6 | ~7 [avatar split clause 1412,1567,1563,1559,1607,1583,1578]
1618. ~1 | ~2 | 11 | 13 | ~5 | ~6 | ~7 [avatar split clause 1411,1567,1563,1559,1607,1591,1547,1543]
1619. 1 | 2 | 11 | 13 | ~5 | ~6 | ~7 [avatar split clause 1410,1567,1563,1559,1607,1591,1547,1543]
1620. 8 | 11 | 13 | ~5 | ~6 | ~7 [avatar split clause 1409,1567,1563,1559,1607,1591,1573]
1621. 9 | 11 | 13 | ~5 | ~6 | ~7 [avatar split clause 1408,1567,1563,1559,1607,1591,1578]
1622. ~1 | ~2 | 12 | 13 | ~5 | ~6 | ~7 [avatar split clause 1407,1567,1563,1559,1607,1599,1547,1543]
1623. 1 | 2 | 12 | 13 | ~5 | ~6 | ~7 [avatar split clause 1406,1567,1563,1559,1607,1599,1547,1543]
1624. 8 | 12 | 13 | ~5 | ~6 | ~7 [avatar split clause 1405,1567,1563,1559,1607,1599,1573]
1625. 9 | 12 | 13 | ~5 | ~6 | ~7 [avatar split clause 1404,1567,1563,1559,1607,1599,1578]
1627. 14 <=> m1_subset_1(sK8,u1_struct_0(sK0)) [avatar definition]
1629. m1_subset_1(sK8,u1_struct_0(sK0)) <- (14) [avatar component clause 1627] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1630. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | ~7 [avatar split clause 1403,1567,1563,1559,1627,1551,1547,1543]
1631. 1 | 2 | ~3 | 14 | ~5 | ~6 | ~7 [avatar split clause 1402,1567,1563,1559,1627,1551,1547,1543]
1632. 8 | ~3 | 14 | ~5 | ~6 | ~7 [avatar split clause 1401,1567,1563,1559,1627,1551,1573]
1633. 9 | ~3 | 14 | ~5 | ~6 | ~7 [avatar split clause 1400,1567,1563,1559,1627,1551,1578]
1634. ~1 | ~2 | 10 | 14 | ~5 | ~6 | ~7 [avatar split clause 1399,1567,1563,1559,1627,1583,1547,1543]
1635. 1 | 2 | 10 | 14 | ~5 | ~6 | ~7 [avatar split clause 1398,1567,1563,1559,1627,1583,1547,1543]
1636. 8 | 10 | 14 | ~5 | ~6 | ~7 [avatar split clause 1397,1567,1563,1559,1627,1583,1573]
1637. 9 | 10 | 14 | ~5 | ~6 | ~7 [avatar split clause 1396,1567,1563,1559,1627,1583,1578]
1638. ~1 | ~2 | 11 | 14 | ~5 | ~6 | ~7 [avatar split clause 1395,1567,1563,1559,1627,1591,1547,1543]
1639. 1 | 2 | 11 | 14 | ~5 | ~6 | ~7 [avatar split clause 1394,1567,1563,1559,1627,1591,1547,1543]
1640. 8 | 11 | 14 | ~5 | ~6 | ~7 [avatar split clause 1393,1567,1563,1559,1627,1591,1573]
1641. 9 | 11 | 14 | ~5 | ~6 | ~7 [avatar split clause 1392,1567,1563,1559,1627,1591,1578]
1642. ~1 | ~2 | 12 | 14 | ~5 | ~6 | ~7 [avatar split clause 1391,1567,1563,1559,1627,1599,1547,1543]
1643. 1 | 2 | 12 | 14 | ~5 | ~6 | ~7 [avatar split clause 1390,1567,1563,1559,1627,1599,1547,1543]
1644. 8 | 12 | 14 | ~5 | ~6 | ~7 [avatar split clause 1389,1567,1563,1559,1627,1599,1573]
1645. 9 | 12 | 14 | ~5 | ~6 | ~7 [avatar split clause 1388,1567,1563,1559,1627,1599,1578]
1647. 15 <=> m1_subset_1(sK7,u1_struct_0(sK0)) [avatar definition]
1650. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | ~7 [avatar split clause 1387,1567,1563,1559,1647,1551,1547,1543]
1651. 1 | 2 | ~3 | 15 | ~5 | ~6 | ~7 [avatar split clause 1386,1567,1563,1559,1647,1551,1547,1543]
1652. 8 | ~3 | 15 | ~5 | ~6 | ~7 [avatar split clause 1385,1567,1563,1559,1647,1551,1573]
1653. 9 | ~3 | 15 | ~5 | ~6 | ~7 [avatar split clause 1384,1567,1563,1559,1647,1551,1578]
1654. ~1 | ~2 | 10 | 15 | ~5 | ~6 | ~7 [avatar split clause 1383,1567,1563,1559,1647,1583,1547,1543]
1655. 1 | 2 | 10 | 15 | ~5 | ~6 | ~7 [avatar split clause 1382,1567,1563,1559,1647,1583,1547,1543]
1656. 8 | 10 | 15 | ~5 | ~6 | ~7 [avatar split clause 1381,1567,1563,1559,1647,1583,1573]
1657. 9 | 10 | 15 | ~5 | ~6 | ~7 [avatar split clause 1380,1567,1563,1559,1647,1583,1578]
1658. ~1 | ~2 | 11 | 15 | ~5 | ~6 | ~7 [avatar split clause 1379,1567,1563,1559,1647,1591,1547,1543]
1659. 1 | 2 | 11 | 15 | ~5 | ~6 | ~7 [avatar split clause 1378,1567,1563,1559,1647,1591,1547,1543]
1660. 8 | 11 | 15 | ~5 | ~6 | ~7 [avatar split clause 1377,1567,1563,1559,1647,1591,1573]
1661. 9 | 11 | 15 | ~5 | ~6 | ~7 [avatar split clause 1376,1567,1563,1559,1647,1591,1578]
1662. ~1 | ~2 | 12 | 15 | ~5 | ~6 | ~7 [avatar split clause 1375,1567,1563,1559,1647,1599,1547,1543]
1663. 1 | 2 | 12 | 15 | ~5 | ~6 | ~7 [avatar split clause 1374,1567,1563,1559,1647,1599,1547,1543]
1664. 8 | 12 | 15 | ~5 | ~6 | ~7 [avatar split clause 1373,1567,1563,1559,1647,1599,1573]
1665. 9 | 12 | 15 | ~5 | ~6 | ~7 [avatar split clause 1372,1567,1563,1559,1647,1599,1578]
1667. 16 <=> m1_subset_1(sK6,u1_struct_0(sK0)) [avatar definition]
1669. m1_subset_1(sK6,u1_struct_0(sK0)) <- (16) [avatar component clause 1667] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1670. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | ~7 [avatar split clause 1371,1567,1563,1559,1667,1551,1547,1543]
1671. 1 | 2 | ~3 | 16 | ~5 | ~6 | ~7 [avatar split clause 1370,1567,1563,1559,1667,1551,1547,1543]
1672. 8 | ~3 | 16 | ~5 | ~6 | ~7 [avatar split clause 1369,1567,1563,1559,1667,1551,1573]
1673. 9 | ~3 | 16 | ~5 | ~6 | ~7 [avatar split clause 1368,1567,1563,1559,1667,1551,1578]
1674. ~1 | ~2 | 10 | 16 | ~5 | ~6 | ~7 [avatar split clause 1367,1567,1563,1559,1667,1583,1547,1543]
1675. 1 | 2 | 10 | 16 | ~5 | ~6 | ~7 [avatar split clause 1366,1567,1563,1559,1667,1583,1547,1543]
1676. 8 | 10 | 16 | ~5 | ~6 | ~7 [avatar split clause 1365,1567,1563,1559,1667,1583,1573]
1677. 9 | 10 | 16 | ~5 | ~6 | ~7 [avatar split clause 1364,1567,1563,1559,1667,1583,1578]
1678. ~1 | ~2 | 11 | 16 | ~5 | ~6 | ~7 [avatar split clause 1363,1567,1563,1559,1667,1591,1547,1543]
1679. 1 | 2 | 11 | 16 | ~5 | ~6 | ~7 [avatar split clause 1362,1567,1563,1559,1667,1591,1547,1543]
1680. 8 | 11 | 16 | ~5 | ~6 | ~7 [avatar split clause 1361,1567,1563,1559,1667,1591,1573]
1681. 9 | 11 | 16 | ~5 | ~6 | ~7 [avatar split clause 1360,1567,1563,1559,1667,1591,1578]
1682. ~1 | ~2 | 12 | 16 | ~5 | ~6 | ~7 [avatar split clause 1359,1567,1563,1559,1667,1599,1547,1543]
1683. 1 | 2 | 12 | 16 | ~5 | ~6 | ~7 [avatar split clause 1358,1567,1563,1559,1667,1599,1547,1543]
1684. 8 | 12 | 16 | ~5 | ~6 | ~7 [avatar split clause 1357,1567,1563,1559,1667,1599,1573]
1685. 9 | 12 | 16 | ~5 | ~6 | ~7 [avatar split clause 1356,1567,1563,1559,1667,1599,1578]
1687. 17 <=> m1_subset_1(sK5,u1_struct_0(sK0)) [avatar definition]
1689. m1_subset_1(sK5,u1_struct_0(sK0)) <- (17) [avatar component clause 1687] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1690. ~1 | ~2 | ~3 | ~4 | 17 | ~7 [avatar split clause 1355,1567,1687,1555,1551,1547,1543]
1691. 1 | 2 | ~3 | ~4 | 17 | ~7 [avatar split clause 1354,1567,1687,1555,1551,1547,1543]
1692. 8 | ~3 | ~4 | 17 | ~7 [avatar split clause 1353,1567,1687,1555,1551,1573]
1693. 9 | ~3 | ~4 | 17 | ~7 [avatar split clause 1352,1567,1687,1555,1551,1578]
1694. ~1 | ~2 | 10 | ~4 | 17 | ~7 [avatar split clause 1351,1567,1687,1555,1583,1547,1543]
1695. 1 | 2 | 10 | ~4 | 17 | ~7 [avatar split clause 1350,1567,1687,1555,1583,1547,1543]
1696. 8 | 10 | ~4 | 17 | ~7 [avatar split clause 1349,1567,1687,1555,1583,1573]
1697. 9 | 10 | ~4 | 17 | ~7 [avatar split clause 1348,1567,1687,1555,1583,1578]
1698. ~1 | ~2 | 11 | ~4 | 17 | ~7 [avatar split clause 1347,1567,1687,1555,1591,1547,1543]
1699. 1 | 2 | 11 | ~4 | 17 | ~7 [avatar split clause 1346,1567,1687,1555,1591,1547,1543]
1700. 8 | 11 | ~4 | 17 | ~7 [avatar split clause 1345,1567,1687,1555,1591,1573]
1701. 9 | 11 | ~4 | 17 | ~7 [avatar split clause 1344,1567,1687,1555,1591,1578]
1702. ~1 | ~2 | 12 | ~4 | 17 | ~7 [avatar split clause 1343,1567,1687,1555,1599,1547,1543]
1703. 1 | 2 | 12 | ~4 | 17 | ~7 [avatar split clause 1342,1567,1687,1555,1599,1547,1543]
1704. 8 | 12 | ~4 | 17 | ~7 [avatar split clause 1341,1567,1687,1555,1599,1573]
1705. 9 | 12 | ~4 | 17 | ~7 [avatar split clause 1340,1567,1687,1555,1599,1578]
1706. ~1 | ~2 | ~3 | 13 | 17 | ~7 [avatar split clause 1339,1567,1687,1607,1551,1547,1543]
1707. 1 | 2 | ~3 | 13 | 17 | ~7 [avatar split clause 1338,1567,1687,1607,1551,1547,1543]
1708. 8 | ~3 | 13 | 17 | ~7 [avatar split clause 1337,1567,1687,1607,1551,1573]
1709. 9 | ~3 | 13 | 17 | ~7 [avatar split clause 1336,1567,1687,1607,1551,1578]
1710. ~1 | ~2 | 10 | 13 | 17 | ~7 [avatar split clause 1335,1567,1687,1607,1583,1547,1543]
1711. 1 | 2 | 10 | 13 | 17 | ~7 [avatar split clause 1334,1567,1687,1607,1583,1547,1543]
1712. 8 | 10 | 13 | 17 | ~7 [avatar split clause 1333,1567,1687,1607,1583,1573]
1713. 9 | 10 | 13 | 17 | ~7 [avatar split clause 1332,1567,1687,1607,1583,1578]
1714. ~1 | ~2 | 11 | 13 | 17 | ~7 [avatar split clause 1331,1567,1687,1607,1591,1547,1543]
1715. 1 | 2 | 11 | 13 | 17 | ~7 [avatar split clause 1330,1567,1687,1607,1591,1547,1543]
1716. 8 | 11 | 13 | 17 | ~7 [avatar split clause 1329,1567,1687,1607,1591,1573]
1717. 9 | 11 | 13 | 17 | ~7 [avatar split clause 1328,1567,1687,1607,1591,1578]
1718. ~1 | ~2 | 12 | 13 | 17 | ~7 [avatar split clause 1327,1567,1687,1607,1599,1547,1543]
1719. 1 | 2 | 12 | 13 | 17 | ~7 [avatar split clause 1326,1567,1687,1607,1599,1547,1543]
1720. 8 | 12 | 13 | 17 | ~7 [avatar split clause 1325,1567,1687,1607,1599,1573]
1721. 9 | 12 | 13 | 17 | ~7 [avatar split clause 1324,1567,1687,1607,1599,1578]
1722. ~1 | ~2 | ~3 | 14 | 17 | ~7 [avatar split clause 1323,1567,1687,1627,1551,1547,1543]
1723. 1 | 2 | ~3 | 14 | 17 | ~7 [avatar split clause 1322,1567,1687,1627,1551,1547,1543]
1724. 8 | ~3 | 14 | 17 | ~7 [avatar split clause 1321,1567,1687,1627,1551,1573]
1725. 9 | ~3 | 14 | 17 | ~7 [avatar split clause 1320,1567,1687,1627,1551,1578]
1726. ~1 | ~2 | 10 | 14 | 17 | ~7 [avatar split clause 1319,1567,1687,1627,1583,1547,1543]
1727. 1 | 2 | 10 | 14 | 17 | ~7 [avatar split clause 1318,1567,1687,1627,1583,1547,1543]
1728. 8 | 10 | 14 | 17 | ~7 [avatar split clause 1317,1567,1687,1627,1583,1573]
1729. 9 | 10 | 14 | 17 | ~7 [avatar split clause 1316,1567,1687,1627,1583,1578]
1730. ~1 | ~2 | 11 | 14 | 17 | ~7 [avatar split clause 1315,1567,1687,1627,1591,1547,1543]
1731. 1 | 2 | 11 | 14 | 17 | ~7 [avatar split clause 1314,1567,1687,1627,1591,1547,1543]
1732. 8 | 11 | 14 | 17 | ~7 [avatar split clause 1313,1567,1687,1627,1591,1573]
1733. 9 | 11 | 14 | 17 | ~7 [avatar split clause 1312,1567,1687,1627,1591,1578]
1734. ~1 | ~2 | 12 | 14 | 17 | ~7 [avatar split clause 1311,1567,1687,1627,1599,1547,1543]
1735. 1 | 2 | 12 | 14 | 17 | ~7 [avatar split clause 1310,1567,1687,1627,1599,1547,1543]
1736. 8 | 12 | 14 | 17 | ~7 [avatar split clause 1309,1567,1687,1627,1599,1573]
1737. 9 | 12 | 14 | 17 | ~7 [avatar split clause 1308,1567,1687,1627,1599,1578]
1738. ~1 | ~2 | ~3 | 15 | 17 | ~7 [avatar split clause 1307,1567,1687,1647,1551,1547,1543]
1739. 1 | 2 | ~3 | 15 | 17 | ~7 [avatar split clause 1306,1567,1687,1647,1551,1547,1543]
1740. 8 | ~3 | 15 | 17 | ~7 [avatar split clause 1305,1567,1687,1647,1551,1573]
1741. 9 | ~3 | 15 | 17 | ~7 [avatar split clause 1304,1567,1687,1647,1551,1578]
1742. ~1 | ~2 | 10 | 15 | 17 | ~7 [avatar split clause 1303,1567,1687,1647,1583,1547,1543]
1743. 1 | 2 | 10 | 15 | 17 | ~7 [avatar split clause 1302,1567,1687,1647,1583,1547,1543]
1744. 8 | 10 | 15 | 17 | ~7 [avatar split clause 1301,1567,1687,1647,1583,1573]
1745. 9 | 10 | 15 | 17 | ~7 [avatar split clause 1300,1567,1687,1647,1583,1578]
1746. ~1 | ~2 | 11 | 15 | 17 | ~7 [avatar split clause 1299,1567,1687,1647,1591,1547,1543]
1747. 1 | 2 | 11 | 15 | 17 | ~7 [avatar split clause 1298,1567,1687,1647,1591,1547,1543]
1748. 8 | 11 | 15 | 17 | ~7 [avatar split clause 1297,1567,1687,1647,1591,1573]
1749. 9 | 11 | 15 | 17 | ~7 [avatar split clause 1296,1567,1687,1647,1591,1578]
1750. ~1 | ~2 | 12 | 15 | 17 | ~7 [avatar split clause 1295,1567,1687,1647,1599,1547,1543]
1751. 1 | 2 | 12 | 15 | 17 | ~7 [avatar split clause 1294,1567,1687,1647,1599,1547,1543]
1752. 8 | 12 | 15 | 17 | ~7 [avatar split clause 1293,1567,1687,1647,1599,1573]
1753. 9 | 12 | 15 | 17 | ~7 [avatar split clause 1292,1567,1687,1647,1599,1578]
1754. ~1 | ~2 | ~3 | 16 | 17 | ~7 [avatar split clause 1291,1567,1687,1667,1551,1547,1543]
1755. 1 | 2 | ~3 | 16 | 17 | ~7 [avatar split clause 1290,1567,1687,1667,1551,1547,1543]
1756. 8 | ~3 | 16 | 17 | ~7 [avatar split clause 1289,1567,1687,1667,1551,1573]
1757. 9 | ~3 | 16 | 17 | ~7 [avatar split clause 1288,1567,1687,1667,1551,1578]
1758. ~1 | ~2 | 10 | 16 | 17 | ~7 [avatar split clause 1287,1567,1687,1667,1583,1547,1543]
1759. 1 | 2 | 10 | 16 | 17 | ~7 [avatar split clause 1286,1567,1687,1667,1583,1547,1543]
1760. 8 | 10 | 16 | 17 | ~7 [avatar split clause 1285,1567,1687,1667,1583,1573]
1761. 9 | 10 | 16 | 17 | ~7 [avatar split clause 1284,1567,1687,1667,1583,1578]
1762. ~1 | ~2 | 11 | 16 | 17 | ~7 [avatar split clause 1283,1567,1687,1667,1591,1547,1543]
1763. 1 | 2 | 11 | 16 | 17 | ~7 [avatar split clause 1282,1567,1687,1667,1591,1547,1543]
1764. 8 | 11 | 16 | 17 | ~7 [avatar split clause 1281,1567,1687,1667,1591,1573]
1765. 9 | 11 | 16 | 17 | ~7 [avatar split clause 1280,1567,1687,1667,1591,1578]
1766. ~1 | ~2 | 12 | 16 | 17 | ~7 [avatar split clause 1279,1567,1687,1667,1599,1547,1543]
1767. 1 | 2 | 12 | 16 | 17 | ~7 [avatar split clause 1278,1567,1687,1667,1599,1547,1543]
1768. 8 | 12 | 16 | 17 | ~7 [avatar split clause 1277,1567,1687,1667,1599,1573]
1769. 9 | 12 | 16 | 17 | ~7 [avatar split clause 1276,1567,1687,1667,1599,1578]
1771. 18 <=> m1_subset_1(sK4,u1_struct_0(sK0)) [avatar definition]
1773. m1_subset_1(sK4,u1_struct_0(sK0)) <- (18) [avatar component clause 1771] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1774. ~1 | ~2 | ~3 | ~4 | 18 | ~7 [avatar split clause 1275,1567,1771,1555,1551,1547,1543]
1775. 1 | 2 | ~3 | ~4 | 18 | ~7 [avatar split clause 1274,1567,1771,1555,1551,1547,1543]
1776. 8 | ~3 | ~4 | 18 | ~7 [avatar split clause 1273,1567,1771,1555,1551,1573]
1777. 9 | ~3 | ~4 | 18 | ~7 [avatar split clause 1272,1567,1771,1555,1551,1578]
1778. ~1 | ~2 | 10 | ~4 | 18 | ~7 [avatar split clause 1271,1567,1771,1555,1583,1547,1543]
1779. 1 | 2 | 10 | ~4 | 18 | ~7 [avatar split clause 1270,1567,1771,1555,1583,1547,1543]
1780. 8 | 10 | ~4 | 18 | ~7 [avatar split clause 1269,1567,1771,1555,1583,1573]
1781. 9 | 10 | ~4 | 18 | ~7 [avatar split clause 1268,1567,1771,1555,1583,1578]
1782. ~1 | ~2 | 11 | ~4 | 18 | ~7 [avatar split clause 1267,1567,1771,1555,1591,1547,1543]
1783. 1 | 2 | 11 | ~4 | 18 | ~7 [avatar split clause 1266,1567,1771,1555,1591,1547,1543]
1784. 8 | 11 | ~4 | 18 | ~7 [avatar split clause 1265,1567,1771,1555,1591,1573]
1785. 9 | 11 | ~4 | 18 | ~7 [avatar split clause 1264,1567,1771,1555,1591,1578]
1786. ~1 | ~2 | 12 | ~4 | 18 | ~7 [avatar split clause 1263,1567,1771,1555,1599,1547,1543]
1787. 1 | 2 | 12 | ~4 | 18 | ~7 [avatar split clause 1262,1567,1771,1555,1599,1547,1543]
1788. 8 | 12 | ~4 | 18 | ~7 [avatar split clause 1261,1567,1771,1555,1599,1573]
1789. 9 | 12 | ~4 | 18 | ~7 [avatar split clause 1260,1567,1771,1555,1599,1578]
1790. ~1 | ~2 | ~3 | 13 | 18 | ~7 [avatar split clause 1259,1567,1771,1607,1551,1547,1543]
1791. 1 | 2 | ~3 | 13 | 18 | ~7 [avatar split clause 1258,1567,1771,1607,1551,1547,1543]
1792. 8 | ~3 | 13 | 18 | ~7 [avatar split clause 1257,1567,1771,1607,1551,1573]
1793. 9 | ~3 | 13 | 18 | ~7 [avatar split clause 1256,1567,1771,1607,1551,1578]
1794. ~1 | ~2 | 10 | 13 | 18 | ~7 [avatar split clause 1255,1567,1771,1607,1583,1547,1543]
1795. 1 | 2 | 10 | 13 | 18 | ~7 [avatar split clause 1254,1567,1771,1607,1583,1547,1543]
1796. 8 | 10 | 13 | 18 | ~7 [avatar split clause 1253,1567,1771,1607,1583,1573]
1797. 9 | 10 | 13 | 18 | ~7 [avatar split clause 1252,1567,1771,1607,1583,1578]
1798. ~1 | ~2 | 11 | 13 | 18 | ~7 [avatar split clause 1251,1567,1771,1607,1591,1547,1543]
1799. 1 | 2 | 11 | 13 | 18 | ~7 [avatar split clause 1250,1567,1771,1607,1591,1547,1543]
1800. 8 | 11 | 13 | 18 | ~7 [avatar split clause 1249,1567,1771,1607,1591,1573]
1801. 9 | 11 | 13 | 18 | ~7 [avatar split clause 1248,1567,1771,1607,1591,1578]
1802. ~1 | ~2 | 12 | 13 | 18 | ~7 [avatar split clause 1247,1567,1771,1607,1599,1547,1543]
1803. 1 | 2 | 12 | 13 | 18 | ~7 [avatar split clause 1246,1567,1771,1607,1599,1547,1543]
1804. 8 | 12 | 13 | 18 | ~7 [avatar split clause 1245,1567,1771,1607,1599,1573]
1805. 9 | 12 | 13 | 18 | ~7 [avatar split clause 1244,1567,1771,1607,1599,1578]
1806. ~1 | ~2 | ~3 | 14 | 18 | ~7 [avatar split clause 1243,1567,1771,1627,1551,1547,1543]
1807. 1 | 2 | ~3 | 14 | 18 | ~7 [avatar split clause 1242,1567,1771,1627,1551,1547,1543]
1808. 8 | ~3 | 14 | 18 | ~7 [avatar split clause 1241,1567,1771,1627,1551,1573]
1809. 9 | ~3 | 14 | 18 | ~7 [avatar split clause 1240,1567,1771,1627,1551,1578]
1810. ~1 | ~2 | 10 | 14 | 18 | ~7 [avatar split clause 1239,1567,1771,1627,1583,1547,1543]
1811. 1 | 2 | 10 | 14 | 18 | ~7 [avatar split clause 1238,1567,1771,1627,1583,1547,1543]
1812. 8 | 10 | 14 | 18 | ~7 [avatar split clause 1237,1567,1771,1627,1583,1573]
1813. 9 | 10 | 14 | 18 | ~7 [avatar split clause 1236,1567,1771,1627,1583,1578]
1814. ~1 | ~2 | 11 | 14 | 18 | ~7 [avatar split clause 1235,1567,1771,1627,1591,1547,1543]
1815. 1 | 2 | 11 | 14 | 18 | ~7 [avatar split clause 1234,1567,1771,1627,1591,1547,1543]
1816. 8 | 11 | 14 | 18 | ~7 [avatar split clause 1233,1567,1771,1627,1591,1573]
1817. 9 | 11 | 14 | 18 | ~7 [avatar split clause 1232,1567,1771,1627,1591,1578]
1818. ~1 | ~2 | 12 | 14 | 18 | ~7 [avatar split clause 1231,1567,1771,1627,1599,1547,1543]
1819. 1 | 2 | 12 | 14 | 18 | ~7 [avatar split clause 1230,1567,1771,1627,1599,1547,1543]
1820. 8 | 12 | 14 | 18 | ~7 [avatar split clause 1229,1567,1771,1627,1599,1573]
1821. 9 | 12 | 14 | 18 | ~7 [avatar split clause 1228,1567,1771,1627,1599,1578]
1822. ~1 | ~2 | ~3 | 15 | 18 | ~7 [avatar split clause 1227,1567,1771,1647,1551,1547,1543]
1823. 1 | 2 | ~3 | 15 | 18 | ~7 [avatar split clause 1226,1567,1771,1647,1551,1547,1543]
1824. 8 | ~3 | 15 | 18 | ~7 [avatar split clause 1225,1567,1771,1647,1551,1573]
1825. 9 | ~3 | 15 | 18 | ~7 [avatar split clause 1224,1567,1771,1647,1551,1578]
1826. ~1 | ~2 | 10 | 15 | 18 | ~7 [avatar split clause 1223,1567,1771,1647,1583,1547,1543]
1827. 1 | 2 | 10 | 15 | 18 | ~7 [avatar split clause 1222,1567,1771,1647,1583,1547,1543]
1828. 8 | 10 | 15 | 18 | ~7 [avatar split clause 1221,1567,1771,1647,1583,1573]
1829. 9 | 10 | 15 | 18 | ~7 [avatar split clause 1220,1567,1771,1647,1583,1578]
1830. ~1 | ~2 | 11 | 15 | 18 | ~7 [avatar split clause 1219,1567,1771,1647,1591,1547,1543]
1831. 1 | 2 | 11 | 15 | 18 | ~7 [avatar split clause 1218,1567,1771,1647,1591,1547,1543]
1832. 8 | 11 | 15 | 18 | ~7 [avatar split clause 1217,1567,1771,1647,1591,1573]
1833. 9 | 11 | 15 | 18 | ~7 [avatar split clause 1216,1567,1771,1647,1591,1578]
1834. ~1 | ~2 | 12 | 15 | 18 | ~7 [avatar split clause 1215,1567,1771,1647,1599,1547,1543]
1835. 1 | 2 | 12 | 15 | 18 | ~7 [avatar split clause 1214,1567,1771,1647,1599,1547,1543]
1836. 8 | 12 | 15 | 18 | ~7 [avatar split clause 1213,1567,1771,1647,1599,1573]
1837. 9 | 12 | 15 | 18 | ~7 [avatar split clause 1212,1567,1771,1647,1599,1578]
1838. ~1 | ~2 | ~3 | 16 | 18 | ~7 [avatar split clause 1211,1567,1771,1667,1551,1547,1543]
1839. 1 | 2 | ~3 | 16 | 18 | ~7 [avatar split clause 1210,1567,1771,1667,1551,1547,1543]
1840. 8 | ~3 | 16 | 18 | ~7 [avatar split clause 1209,1567,1771,1667,1551,1573]
1841. 9 | ~3 | 16 | 18 | ~7 [avatar split clause 1208,1567,1771,1667,1551,1578]
1842. ~1 | ~2 | 10 | 16 | 18 | ~7 [avatar split clause 1207,1567,1771,1667,1583,1547,1543]
1843. 1 | 2 | 10 | 16 | 18 | ~7 [avatar split clause 1206,1567,1771,1667,1583,1547,1543]
1844. 8 | 10 | 16 | 18 | ~7 [avatar split clause 1205,1567,1771,1667,1583,1573]
1845. 9 | 10 | 16 | 18 | ~7 [avatar split clause 1204,1567,1771,1667,1583,1578]
1846. ~1 | ~2 | 11 | 16 | 18 | ~7 [avatar split clause 1203,1567,1771,1667,1591,1547,1543]
1847. 1 | 2 | 11 | 16 | 18 | ~7 [avatar split clause 1202,1567,1771,1667,1591,1547,1543]
1848. 8 | 11 | 16 | 18 | ~7 [avatar split clause 1201,1567,1771,1667,1591,1573]
1849. 9 | 11 | 16 | 18 | ~7 [avatar split clause 1200,1567,1771,1667,1591,1578]
1850. ~1 | ~2 | 12 | 16 | 18 | ~7 [avatar split clause 1199,1567,1771,1667,1599,1547,1543]
1851. 1 | 2 | 12 | 16 | 18 | ~7 [avatar split clause 1198,1567,1771,1667,1599,1547,1543]
1852. 8 | 12 | 16 | 18 | ~7 [avatar split clause 1197,1567,1771,1667,1599,1573]
1853. 9 | 12 | 16 | 18 | ~7 [avatar split clause 1196,1567,1771,1667,1599,1578]
1855. 19 <=> r1_bcialg_1(sK0,sK1,sK3) [avatar definition]
1857. r1_bcialg_1(sK0,sK1,sK3) <- (19) [avatar component clause 1855] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1858. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | 19 [avatar split clause 1195,1855,1563,1559,1555,1551,1547,1543]
1859. 1 | 2 | ~3 | ~4 | ~5 | ~6 | 19 [avatar split clause 1194,1855,1563,1559,1555,1551,1547,1543]
1860. 8 | ~3 | ~4 | ~5 | ~6 | 19 [avatar split clause 1193,1855,1563,1559,1555,1551,1573]
1861. 9 | ~3 | ~4 | ~5 | ~6 | 19 [avatar split clause 1192,1855,1563,1559,1555,1551,1578]
1862. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | 19 [avatar split clause 1191,1855,1563,1559,1555,1583,1547,1543]
1863. 1 | 2 | 10 | ~4 | ~5 | ~6 | 19 [avatar split clause 1190,1855,1563,1559,1555,1583,1547,1543]
1864. 8 | 10 | ~4 | ~5 | ~6 | 19 [avatar split clause 1189,1855,1563,1559,1555,1583,1573]
1865. 9 | 10 | ~4 | ~5 | ~6 | 19 [avatar split clause 1188,1855,1563,1559,1555,1583,1578]
1866. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | 19 [avatar split clause 1187,1855,1563,1559,1555,1591,1547,1543]
1867. 1 | 2 | 11 | ~4 | ~5 | ~6 | 19 [avatar split clause 1186,1855,1563,1559,1555,1591,1547,1543]
1868. 8 | 11 | ~4 | ~5 | ~6 | 19 [avatar split clause 1185,1855,1563,1559,1555,1591,1573]
1869. 9 | 11 | ~4 | ~5 | ~6 | 19 [avatar split clause 1184,1855,1563,1559,1555,1591,1578]
1870. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | 19 [avatar split clause 1183,1855,1563,1559,1555,1599,1547,1543]
1871. 1 | 2 | 12 | ~4 | ~5 | ~6 | 19 [avatar split clause 1182,1855,1563,1559,1555,1599,1547,1543]
1872. 8 | 12 | ~4 | ~5 | ~6 | 19 [avatar split clause 1181,1855,1563,1559,1555,1599,1573]
1873. 9 | 12 | ~4 | ~5 | ~6 | 19 [avatar split clause 1180,1855,1563,1559,1555,1599,1578]
1874. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | 19 [avatar split clause 1179,1855,1563,1559,1607,1551,1547,1543]
1875. 1 | 2 | ~3 | 13 | ~5 | ~6 | 19 [avatar split clause 1178,1855,1563,1559,1607,1551,1547,1543]
1876. 8 | ~3 | 13 | ~5 | ~6 | 19 [avatar split clause 1177,1855,1563,1559,1607,1551,1573]
1877. 9 | ~3 | 13 | ~5 | ~6 | 19 [avatar split clause 1176,1855,1563,1559,1607,1551,1578]
1878. ~1 | ~2 | 10 | 13 | ~5 | ~6 | 19 [avatar split clause 1175,1855,1563,1559,1607,1583,1547,1543]
1879. 1 | 2 | 10 | 13 | ~5 | ~6 | 19 [avatar split clause 1174,1855,1563,1559,1607,1583,1547,1543]
1880. 8 | 10 | 13 | ~5 | ~6 | 19 [avatar split clause 1173,1855,1563,1559,1607,1583,1573]
1881. 9 | 10 | 13 | ~5 | ~6 | 19 [avatar split clause 1172,1855,1563,1559,1607,1583,1578]
1882. ~1 | ~2 | 11 | 13 | ~5 | ~6 | 19 [avatar split clause 1171,1855,1563,1559,1607,1591,1547,1543]
1883. 1 | 2 | 11 | 13 | ~5 | ~6 | 19 [avatar split clause 1170,1855,1563,1559,1607,1591,1547,1543]
1884. 8 | 11 | 13 | ~5 | ~6 | 19 [avatar split clause 1169,1855,1563,1559,1607,1591,1573]
1885. 9 | 11 | 13 | ~5 | ~6 | 19 [avatar split clause 1168,1855,1563,1559,1607,1591,1578]
1886. ~1 | ~2 | 12 | 13 | ~5 | ~6 | 19 [avatar split clause 1167,1855,1563,1559,1607,1599,1547,1543]
1887. 1 | 2 | 12 | 13 | ~5 | ~6 | 19 [avatar split clause 1166,1855,1563,1559,1607,1599,1547,1543]
1888. 8 | 12 | 13 | ~5 | ~6 | 19 [avatar split clause 1165,1855,1563,1559,1607,1599,1573]
1889. 9 | 12 | 13 | ~5 | ~6 | 19 [avatar split clause 1164,1855,1563,1559,1607,1599,1578]
1890. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | 19 [avatar split clause 1163,1855,1563,1559,1627,1551,1547,1543]
1891. 1 | 2 | ~3 | 14 | ~5 | ~6 | 19 [avatar split clause 1162,1855,1563,1559,1627,1551,1547,1543]
1892. 8 | ~3 | 14 | ~5 | ~6 | 19 [avatar split clause 1161,1855,1563,1559,1627,1551,1573]
1893. 9 | ~3 | 14 | ~5 | ~6 | 19 [avatar split clause 1160,1855,1563,1559,1627,1551,1578]
1894. ~1 | ~2 | 10 | 14 | ~5 | ~6 | 19 [avatar split clause 1159,1855,1563,1559,1627,1583,1547,1543]
1895. 1 | 2 | 10 | 14 | ~5 | ~6 | 19 [avatar split clause 1158,1855,1563,1559,1627,1583,1547,1543]
1896. 8 | 10 | 14 | ~5 | ~6 | 19 [avatar split clause 1157,1855,1563,1559,1627,1583,1573]
1897. 9 | 10 | 14 | ~5 | ~6 | 19 [avatar split clause 1156,1855,1563,1559,1627,1583,1578]
1898. ~1 | ~2 | 11 | 14 | ~5 | ~6 | 19 [avatar split clause 1155,1855,1563,1559,1627,1591,1547,1543]
1899. 1 | 2 | 11 | 14 | ~5 | ~6 | 19 [avatar split clause 1154,1855,1563,1559,1627,1591,1547,1543]
1900. 8 | 11 | 14 | ~5 | ~6 | 19 [avatar split clause 1153,1855,1563,1559,1627,1591,1573]
1901. 9 | 11 | 14 | ~5 | ~6 | 19 [avatar split clause 1152,1855,1563,1559,1627,1591,1578]
1902. ~1 | ~2 | 12 | 14 | ~5 | ~6 | 19 [avatar split clause 1151,1855,1563,1559,1627,1599,1547,1543]
1903. 1 | 2 | 12 | 14 | ~5 | ~6 | 19 [avatar split clause 1150,1855,1563,1559,1627,1599,1547,1543]
1904. 8 | 12 | 14 | ~5 | ~6 | 19 [avatar split clause 1149,1855,1563,1559,1627,1599,1573]
1905. 9 | 12 | 14 | ~5 | ~6 | 19 [avatar split clause 1148,1855,1563,1559,1627,1599,1578]
1906. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | 19 [avatar split clause 1147,1855,1563,1559,1647,1551,1547,1543]
1907. 1 | 2 | ~3 | 15 | ~5 | ~6 | 19 [avatar split clause 1146,1855,1563,1559,1647,1551,1547,1543]
1908. 8 | ~3 | 15 | ~5 | ~6 | 19 [avatar split clause 1145,1855,1563,1559,1647,1551,1573]
1909. 9 | ~3 | 15 | ~5 | ~6 | 19 [avatar split clause 1144,1855,1563,1559,1647,1551,1578]
1910. ~1 | ~2 | 10 | 15 | ~5 | ~6 | 19 [avatar split clause 1143,1855,1563,1559,1647,1583,1547,1543]
1911. 1 | 2 | 10 | 15 | ~5 | ~6 | 19 [avatar split clause 1142,1855,1563,1559,1647,1583,1547,1543]
1912. 8 | 10 | 15 | ~5 | ~6 | 19 [avatar split clause 1141,1855,1563,1559,1647,1583,1573]
1913. 9 | 10 | 15 | ~5 | ~6 | 19 [avatar split clause 1140,1855,1563,1559,1647,1583,1578]
1914. ~1 | ~2 | 11 | 15 | ~5 | ~6 | 19 [avatar split clause 1139,1855,1563,1559,1647,1591,1547,1543]
1915. 1 | 2 | 11 | 15 | ~5 | ~6 | 19 [avatar split clause 1138,1855,1563,1559,1647,1591,1547,1543]
1916. 8 | 11 | 15 | ~5 | ~6 | 19 [avatar split clause 1137,1855,1563,1559,1647,1591,1573]
1917. 9 | 11 | 15 | ~5 | ~6 | 19 [avatar split clause 1136,1855,1563,1559,1647,1591,1578]
1918. ~1 | ~2 | 12 | 15 | ~5 | ~6 | 19 [avatar split clause 1135,1855,1563,1559,1647,1599,1547,1543]
1919. 1 | 2 | 12 | 15 | ~5 | ~6 | 19 [avatar split clause 1134,1855,1563,1559,1647,1599,1547,1543]
1920. 8 | 12 | 15 | ~5 | ~6 | 19 [avatar split clause 1133,1855,1563,1559,1647,1599,1573]
1921. 9 | 12 | 15 | ~5 | ~6 | 19 [avatar split clause 1132,1855,1563,1559,1647,1599,1578]
1922. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | 19 [avatar split clause 1131,1855,1563,1559,1667,1551,1547,1543]
1923. 1 | 2 | ~3 | 16 | ~5 | ~6 | 19 [avatar split clause 1130,1855,1563,1559,1667,1551,1547,1543]
1924. 8 | ~3 | 16 | ~5 | ~6 | 19 [avatar split clause 1129,1855,1563,1559,1667,1551,1573]
1925. 9 | ~3 | 16 | ~5 | ~6 | 19 [avatar split clause 1128,1855,1563,1559,1667,1551,1578]
1926. ~1 | ~2 | 10 | 16 | ~5 | ~6 | 19 [avatar split clause 1127,1855,1563,1559,1667,1583,1547,1543]
1927. 1 | 2 | 10 | 16 | ~5 | ~6 | 19 [avatar split clause 1126,1855,1563,1559,1667,1583,1547,1543]
1928. 8 | 10 | 16 | ~5 | ~6 | 19 [avatar split clause 1125,1855,1563,1559,1667,1583,1573]
1929. 9 | 10 | 16 | ~5 | ~6 | 19 [avatar split clause 1124,1855,1563,1559,1667,1583,1578]
1930. ~1 | ~2 | 11 | 16 | ~5 | ~6 | 19 [avatar split clause 1123,1855,1563,1559,1667,1591,1547,1543]
1931. 1 | 2 | 11 | 16 | ~5 | ~6 | 19 [avatar split clause 1122,1855,1563,1559,1667,1591,1547,1543]
1932. 8 | 11 | 16 | ~5 | ~6 | 19 [avatar split clause 1121,1855,1563,1559,1667,1591,1573]
1933. 9 | 11 | 16 | ~5 | ~6 | 19 [avatar split clause 1120,1855,1563,1559,1667,1591,1578]
1934. ~1 | ~2 | 12 | 16 | ~5 | ~6 | 19 [avatar split clause 1119,1855,1563,1559,1667,1599,1547,1543]
1935. 1 | 2 | 12 | 16 | ~5 | ~6 | 19 [avatar split clause 1118,1855,1563,1559,1667,1599,1547,1543]
1936. 8 | 12 | 16 | ~5 | ~6 | 19 [avatar split clause 1117,1855,1563,1559,1667,1599,1573]
1937. 9 | 12 | 16 | ~5 | ~6 | 19 [avatar split clause 1116,1855,1563,1559,1667,1599,1578]
1938. ~1 | ~2 | ~3 | ~4 | 17 | 19 [avatar split clause 1115,1855,1687,1555,1551,1547,1543]
1939. 1 | 2 | ~3 | ~4 | 17 | 19 [avatar split clause 1114,1855,1687,1555,1551,1547,1543]
1940. 8 | ~3 | ~4 | 17 | 19 [avatar split clause 1113,1855,1687,1555,1551,1573]
1941. 9 | ~3 | ~4 | 17 | 19 [avatar split clause 1112,1855,1687,1555,1551,1578]
1942. ~1 | ~2 | 10 | ~4 | 17 | 19 [avatar split clause 1111,1855,1687,1555,1583,1547,1543]
1943. 1 | 2 | 10 | ~4 | 17 | 19 [avatar split clause 1110,1855,1687,1555,1583,1547,1543]
1944. 8 | 10 | ~4 | 17 | 19 [avatar split clause 1109,1855,1687,1555,1583,1573]
1945. 9 | 10 | ~4 | 17 | 19 [avatar split clause 1108,1855,1687,1555,1583,1578]
1946. ~1 | ~2 | 11 | ~4 | 17 | 19 [avatar split clause 1107,1855,1687,1555,1591,1547,1543]
1947. 1 | 2 | 11 | ~4 | 17 | 19 [avatar split clause 1106,1855,1687,1555,1591,1547,1543]
1948. 8 | 11 | ~4 | 17 | 19 [avatar split clause 1105,1855,1687,1555,1591,1573]
1949. 9 | 11 | ~4 | 17 | 19 [avatar split clause 1104,1855,1687,1555,1591,1578]
1950. ~1 | ~2 | 12 | ~4 | 17 | 19 [avatar split clause 1103,1855,1687,1555,1599,1547,1543]
1951. 1 | 2 | 12 | ~4 | 17 | 19 [avatar split clause 1102,1855,1687,1555,1599,1547,1543]
1952. 8 | 12 | ~4 | 17 | 19 [avatar split clause 1101,1855,1687,1555,1599,1573]
1953. 9 | 12 | ~4 | 17 | 19 [avatar split clause 1100,1855,1687,1555,1599,1578]
1954. ~1 | ~2 | ~3 | 13 | 17 | 19 [avatar split clause 1099,1855,1687,1607,1551,1547,1543]
1955. 1 | 2 | ~3 | 13 | 17 | 19 [avatar split clause 1098,1855,1687,1607,1551,1547,1543]
1956. 8 | ~3 | 13 | 17 | 19 [avatar split clause 1097,1855,1687,1607,1551,1573]
1957. 9 | ~3 | 13 | 17 | 19 [avatar split clause 1096,1855,1687,1607,1551,1578]
1958. ~1 | ~2 | 10 | 13 | 17 | 19 [avatar split clause 1095,1855,1687,1607,1583,1547,1543]
1959. 1 | 2 | 10 | 13 | 17 | 19 [avatar split clause 1094,1855,1687,1607,1583,1547,1543]
1960. 8 | 10 | 13 | 17 | 19 [avatar split clause 1093,1855,1687,1607,1583,1573]
1961. 9 | 10 | 13 | 17 | 19 [avatar split clause 1092,1855,1687,1607,1583,1578]
1962. ~1 | ~2 | 11 | 13 | 17 | 19 [avatar split clause 1091,1855,1687,1607,1591,1547,1543]
1963. 1 | 2 | 11 | 13 | 17 | 19 [avatar split clause 1090,1855,1687,1607,1591,1547,1543]
1964. 8 | 11 | 13 | 17 | 19 [avatar split clause 1089,1855,1687,1607,1591,1573]
1965. 9 | 11 | 13 | 17 | 19 [avatar split clause 1088,1855,1687,1607,1591,1578]
1966. ~1 | ~2 | 12 | 13 | 17 | 19 [avatar split clause 1087,1855,1687,1607,1599,1547,1543]
1967. 1 | 2 | 12 | 13 | 17 | 19 [avatar split clause 1086,1855,1687,1607,1599,1547,1543]
1968. 8 | 12 | 13 | 17 | 19 [avatar split clause 1085,1855,1687,1607,1599,1573]
1969. 9 | 12 | 13 | 17 | 19 [avatar split clause 1084,1855,1687,1607,1599,1578]
1970. ~1 | ~2 | ~3 | 14 | 17 | 19 [avatar split clause 1083,1855,1687,1627,1551,1547,1543]
1971. 1 | 2 | ~3 | 14 | 17 | 19 [avatar split clause 1082,1855,1687,1627,1551,1547,1543]
1972. 8 | ~3 | 14 | 17 | 19 [avatar split clause 1081,1855,1687,1627,1551,1573]
1973. 9 | ~3 | 14 | 17 | 19 [avatar split clause 1080,1855,1687,1627,1551,1578]
1974. ~1 | ~2 | 10 | 14 | 17 | 19 [avatar split clause 1079,1855,1687,1627,1583,1547,1543]
1975. 1 | 2 | 10 | 14 | 17 | 19 [avatar split clause 1078,1855,1687,1627,1583,1547,1543]
1976. 8 | 10 | 14 | 17 | 19 [avatar split clause 1077,1855,1687,1627,1583,1573]
1977. 9 | 10 | 14 | 17 | 19 [avatar split clause 1076,1855,1687,1627,1583,1578]
1978. ~1 | ~2 | 11 | 14 | 17 | 19 [avatar split clause 1075,1855,1687,1627,1591,1547,1543]
1979. 1 | 2 | 11 | 14 | 17 | 19 [avatar split clause 1074,1855,1687,1627,1591,1547,1543]
1980. 8 | 11 | 14 | 17 | 19 [avatar split clause 1073,1855,1687,1627,1591,1573]
1981. 9 | 11 | 14 | 17 | 19 [avatar split clause 1072,1855,1687,1627,1591,1578]
1982. ~1 | ~2 | 12 | 14 | 17 | 19 [avatar split clause 1071,1855,1687,1627,1599,1547,1543]
1983. 1 | 2 | 12 | 14 | 17 | 19 [avatar split clause 1070,1855,1687,1627,1599,1547,1543]
1984. 8 | 12 | 14 | 17 | 19 [avatar split clause 1069,1855,1687,1627,1599,1573]
1985. 9 | 12 | 14 | 17 | 19 [avatar split clause 1068,1855,1687,1627,1599,1578]
1986. ~1 | ~2 | ~3 | 15 | 17 | 19 [avatar split clause 1067,1855,1687,1647,1551,1547,1543]
1987. 1 | 2 | ~3 | 15 | 17 | 19 [avatar split clause 1066,1855,1687,1647,1551,1547,1543]
1988. 8 | ~3 | 15 | 17 | 19 [avatar split clause 1065,1855,1687,1647,1551,1573]
1989. 9 | ~3 | 15 | 17 | 19 [avatar split clause 1064,1855,1687,1647,1551,1578]
1990. ~1 | ~2 | 10 | 15 | 17 | 19 [avatar split clause 1063,1855,1687,1647,1583,1547,1543]
1991. 1 | 2 | 10 | 15 | 17 | 19 [avatar split clause 1062,1855,1687,1647,1583,1547,1543]
1992. 8 | 10 | 15 | 17 | 19 [avatar split clause 1061,1855,1687,1647,1583,1573]
1993. 9 | 10 | 15 | 17 | 19 [avatar split clause 1060,1855,1687,1647,1583,1578]
1994. ~1 | ~2 | 11 | 15 | 17 | 19 [avatar split clause 1059,1855,1687,1647,1591,1547,1543]
1995. 1 | 2 | 11 | 15 | 17 | 19 [avatar split clause 1058,1855,1687,1647,1591,1547,1543]
1996. 8 | 11 | 15 | 17 | 19 [avatar split clause 1057,1855,1687,1647,1591,1573]
1997. 9 | 11 | 15 | 17 | 19 [avatar split clause 1056,1855,1687,1647,1591,1578]
1998. ~1 | ~2 | 12 | 15 | 17 | 19 [avatar split clause 1055,1855,1687,1647,1599,1547,1543]
1999. 1 | 2 | 12 | 15 | 17 | 19 [avatar split clause 1054,1855,1687,1647,1599,1547,1543]
2000. 8 | 12 | 15 | 17 | 19 [avatar split clause 1053,1855,1687,1647,1599,1573]
2001. 9 | 12 | 15 | 17 | 19 [avatar split clause 1052,1855,1687,1647,1599,1578]
2002. ~1 | ~2 | ~3 | 16 | 17 | 19 [avatar split clause 1051,1855,1687,1667,1551,1547,1543]
2003. 1 | 2 | ~3 | 16 | 17 | 19 [avatar split clause 1050,1855,1687,1667,1551,1547,1543]
2004. 8 | ~3 | 16 | 17 | 19 [avatar split clause 1049,1855,1687,1667,1551,1573]
2005. 9 | ~3 | 16 | 17 | 19 [avatar split clause 1048,1855,1687,1667,1551,1578]
2006. ~1 | ~2 | 10 | 16 | 17 | 19 [avatar split clause 1047,1855,1687,1667,1583,1547,1543]
2007. 1 | 2 | 10 | 16 | 17 | 19 [avatar split clause 1046,1855,1687,1667,1583,1547,1543]
2008. 8 | 10 | 16 | 17 | 19 [avatar split clause 1045,1855,1687,1667,1583,1573]
2009. 9 | 10 | 16 | 17 | 19 [avatar split clause 1044,1855,1687,1667,1583,1578]
2010. ~1 | ~2 | 11 | 16 | 17 | 19 [avatar split clause 1043,1855,1687,1667,1591,1547,1543]
2011. 1 | 2 | 11 | 16 | 17 | 19 [avatar split clause 1042,1855,1687,1667,1591,1547,1543]
2012. 8 | 11 | 16 | 17 | 19 [avatar split clause 1041,1855,1687,1667,1591,1573]
2013. 9 | 11 | 16 | 17 | 19 [avatar split clause 1040,1855,1687,1667,1591,1578]
2014. ~1 | ~2 | 12 | 16 | 17 | 19 [avatar split clause 1039,1855,1687,1667,1599,1547,1543]
2015. 1 | 2 | 12 | 16 | 17 | 19 [avatar split clause 1038,1855,1687,1667,1599,1547,1543]
2016. 8 | 12 | 16 | 17 | 19 [avatar split clause 1037,1855,1687,1667,1599,1573]
2017. 9 | 12 | 16 | 17 | 19 [avatar split clause 1036,1855,1687,1667,1599,1578]
2018. ~1 | ~2 | ~3 | ~4 | 18 | 19 [avatar split clause 1035,1855,1771,1555,1551,1547,1543]
2019. 1 | 2 | ~3 | ~4 | 18 | 19 [avatar split clause 1034,1855,1771,1555,1551,1547,1543]
2020. 8 | ~3 | ~4 | 18 | 19 [avatar split clause 1033,1855,1771,1555,1551,1573]
2021. 9 | ~3 | ~4 | 18 | 19 [avatar split clause 1032,1855,1771,1555,1551,1578]
2022. ~1 | ~2 | 10 | ~4 | 18 | 19 [avatar split clause 1031,1855,1771,1555,1583,1547,1543]
2023. 1 | 2 | 10 | ~4 | 18 | 19 [avatar split clause 1030,1855,1771,1555,1583,1547,1543]
2024. 8 | 10 | ~4 | 18 | 19 [avatar split clause 1029,1855,1771,1555,1583,1573]
2025. 9 | 10 | ~4 | 18 | 19 [avatar split clause 1028,1855,1771,1555,1583,1578]
2026. ~1 | ~2 | 11 | ~4 | 18 | 19 [avatar split clause 1027,1855,1771,1555,1591,1547,1543]
2027. 1 | 2 | 11 | ~4 | 18 | 19 [avatar split clause 1026,1855,1771,1555,1591,1547,1543]
2028. 8 | 11 | ~4 | 18 | 19 [avatar split clause 1025,1855,1771,1555,1591,1573]
2029. 9 | 11 | ~4 | 18 | 19 [avatar split clause 1024,1855,1771,1555,1591,1578]
2030. ~1 | ~2 | 12 | ~4 | 18 | 19 [avatar split clause 1023,1855,1771,1555,1599,1547,1543]
2031. 1 | 2 | 12 | ~4 | 18 | 19 [avatar split clause 1022,1855,1771,1555,1599,1547,1543]
2032. 8 | 12 | ~4 | 18 | 19 [avatar split clause 1021,1855,1771,1555,1599,1573]
2033. 9 | 12 | ~4 | 18 | 19 [avatar split clause 1020,1855,1771,1555,1599,1578]
2034. ~1 | ~2 | ~3 | 13 | 18 | 19 [avatar split clause 1019,1855,1771,1607,1551,1547,1543]
2035. 1 | 2 | ~3 | 13 | 18 | 19 [avatar split clause 1018,1855,1771,1607,1551,1547,1543]
2036. 8 | ~3 | 13 | 18 | 19 [avatar split clause 1017,1855,1771,1607,1551,1573]
2037. 9 | ~3 | 13 | 18 | 19 [avatar split clause 1016,1855,1771,1607,1551,1578]
2038. ~1 | ~2 | 10 | 13 | 18 | 19 [avatar split clause 1015,1855,1771,1607,1583,1547,1543]
2039. 1 | 2 | 10 | 13 | 18 | 19 [avatar split clause 1014,1855,1771,1607,1583,1547,1543]
2040. 8 | 10 | 13 | 18 | 19 [avatar split clause 1013,1855,1771,1607,1583,1573]
2041. 9 | 10 | 13 | 18 | 19 [avatar split clause 1012,1855,1771,1607,1583,1578]
2042. ~1 | ~2 | 11 | 13 | 18 | 19 [avatar split clause 1011,1855,1771,1607,1591,1547,1543]
2043. 1 | 2 | 11 | 13 | 18 | 19 [avatar split clause 1010,1855,1771,1607,1591,1547,1543]
2044. 8 | 11 | 13 | 18 | 19 [avatar split clause 1009,1855,1771,1607,1591,1573]
2045. 9 | 11 | 13 | 18 | 19 [avatar split clause 1008,1855,1771,1607,1591,1578]
2046. ~1 | ~2 | 12 | 13 | 18 | 19 [avatar split clause 1007,1855,1771,1607,1599,1547,1543]
2047. 1 | 2 | 12 | 13 | 18 | 19 [avatar split clause 1006,1855,1771,1607,1599,1547,1543]
2048. 8 | 12 | 13 | 18 | 19 [avatar split clause 1005,1855,1771,1607,1599,1573]
2049. 9 | 12 | 13 | 18 | 19 [avatar split clause 1004,1855,1771,1607,1599,1578]
2050. ~1 | ~2 | ~3 | 14 | 18 | 19 [avatar split clause 1003,1855,1771,1627,1551,1547,1543]
2051. 1 | 2 | ~3 | 14 | 18 | 19 [avatar split clause 1002,1855,1771,1627,1551,1547,1543]
2052. 8 | ~3 | 14 | 18 | 19 [avatar split clause 1001,1855,1771,1627,1551,1573]
2053. 9 | ~3 | 14 | 18 | 19 [avatar split clause 1000,1855,1771,1627,1551,1578]
2054. ~1 | ~2 | 10 | 14 | 18 | 19 [avatar split clause 999,1855,1771,1627,1583,1547,1543]
2055. 1 | 2 | 10 | 14 | 18 | 19 [avatar split clause 998,1855,1771,1627,1583,1547,1543]
2056. 8 | 10 | 14 | 18 | 19 [avatar split clause 997,1855,1771,1627,1583,1573]
2057. 9 | 10 | 14 | 18 | 19 [avatar split clause 996,1855,1771,1627,1583,1578]
2058. ~1 | ~2 | 11 | 14 | 18 | 19 [avatar split clause 995,1855,1771,1627,1591,1547,1543]
2059. 1 | 2 | 11 | 14 | 18 | 19 [avatar split clause 994,1855,1771,1627,1591,1547,1543]
2060. 8 | 11 | 14 | 18 | 19 [avatar split clause 993,1855,1771,1627,1591,1573]
2061. 9 | 11 | 14 | 18 | 19 [avatar split clause 992,1855,1771,1627,1591,1578]
2062. ~1 | ~2 | 12 | 14 | 18 | 19 [avatar split clause 991,1855,1771,1627,1599,1547,1543]
2063. 1 | 2 | 12 | 14 | 18 | 19 [avatar split clause 990,1855,1771,1627,1599,1547,1543]
2064. 8 | 12 | 14 | 18 | 19 [avatar split clause 989,1855,1771,1627,1599,1573]
2065. 9 | 12 | 14 | 18 | 19 [avatar split clause 988,1855,1771,1627,1599,1578]
2066. ~1 | ~2 | ~3 | 15 | 18 | 19 [avatar split clause 987,1855,1771,1647,1551,1547,1543]
2067. 1 | 2 | ~3 | 15 | 18 | 19 [avatar split clause 986,1855,1771,1647,1551,1547,1543]
2068. 8 | ~3 | 15 | 18 | 19 [avatar split clause 985,1855,1771,1647,1551,1573]
2069. 9 | ~3 | 15 | 18 | 19 [avatar split clause 984,1855,1771,1647,1551,1578]
2070. ~1 | ~2 | 10 | 15 | 18 | 19 [avatar split clause 983,1855,1771,1647,1583,1547,1543]
2071. 1 | 2 | 10 | 15 | 18 | 19 [avatar split clause 982,1855,1771,1647,1583,1547,1543]
2072. 8 | 10 | 15 | 18 | 19 [avatar split clause 981,1855,1771,1647,1583,1573]
2073. 9 | 10 | 15 | 18 | 19 [avatar split clause 980,1855,1771,1647,1583,1578]
2074. ~1 | ~2 | 11 | 15 | 18 | 19 [avatar split clause 979,1855,1771,1647,1591,1547,1543]
2075. 1 | 2 | 11 | 15 | 18 | 19 [avatar split clause 978,1855,1771,1647,1591,1547,1543]
2076. 8 | 11 | 15 | 18 | 19 [avatar split clause 977,1855,1771,1647,1591,1573]
2077. 9 | 11 | 15 | 18 | 19 [avatar split clause 976,1855,1771,1647,1591,1578]
2078. ~1 | ~2 | 12 | 15 | 18 | 19 [avatar split clause 975,1855,1771,1647,1599,1547,1543]
2079. 1 | 2 | 12 | 15 | 18 | 19 [avatar split clause 974,1855,1771,1647,1599,1547,1543]
2080. 8 | 12 | 15 | 18 | 19 [avatar split clause 973,1855,1771,1647,1599,1573]
2081. 9 | 12 | 15 | 18 | 19 [avatar split clause 972,1855,1771,1647,1599,1578]
2082. ~1 | ~2 | ~3 | 16 | 18 | 19 [avatar split clause 971,1855,1771,1667,1551,1547,1543]
2083. 1 | 2 | ~3 | 16 | 18 | 19 [avatar split clause 970,1855,1771,1667,1551,1547,1543]
2084. 8 | ~3 | 16 | 18 | 19 [avatar split clause 969,1855,1771,1667,1551,1573]
2085. 9 | ~3 | 16 | 18 | 19 [avatar split clause 968,1855,1771,1667,1551,1578]
2086. ~1 | ~2 | 10 | 16 | 18 | 19 [avatar split clause 967,1855,1771,1667,1583,1547,1543]
2087. 1 | 2 | 10 | 16 | 18 | 19 [avatar split clause 966,1855,1771,1667,1583,1547,1543]
2088. 8 | 10 | 16 | 18 | 19 [avatar split clause 965,1855,1771,1667,1583,1573]
2089. 9 | 10 | 16 | 18 | 19 [avatar split clause 964,1855,1771,1667,1583,1578]
2090. ~1 | ~2 | 11 | 16 | 18 | 19 [avatar split clause 963,1855,1771,1667,1591,1547,1543]
2091. 1 | 2 | 11 | 16 | 18 | 19 [avatar split clause 962,1855,1771,1667,1591,1547,1543]
2092. 8 | 11 | 16 | 18 | 19 [avatar split clause 961,1855,1771,1667,1591,1573]
2093. 9 | 11 | 16 | 18 | 19 [avatar split clause 960,1855,1771,1667,1591,1578]
2094. ~1 | ~2 | 12 | 16 | 18 | 19 [avatar split clause 959,1855,1771,1667,1599,1547,1543]
2095. 1 | 2 | 12 | 16 | 18 | 19 [avatar split clause 958,1855,1771,1667,1599,1547,1543]
2096. 8 | 12 | 16 | 18 | 19 [avatar split clause 957,1855,1771,1667,1599,1573]
2097. 9 | 12 | 16 | 18 | 19 [avatar split clause 956,1855,1771,1667,1599,1578]
2099. 20 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
2101. m1_subset_1(sK3,u1_struct_0(sK0)) <- (20) [avatar component clause 2099] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2102. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | 20 [avatar split clause 955,2099,1563,1559,1555,1551,1547,1543]
2103. 1 | 2 | ~3 | ~4 | ~5 | ~6 | 20 [avatar split clause 954,2099,1563,1559,1555,1551,1547,1543]
2104. 8 | ~3 | ~4 | ~5 | ~6 | 20 [avatar split clause 953,2099,1563,1559,1555,1551,1573]
2105. 9 | ~3 | ~4 | ~5 | ~6 | 20 [avatar split clause 952,2099,1563,1559,1555,1551,1578]
2106. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | 20 [avatar split clause 951,2099,1563,1559,1555,1583,1547,1543]
2107. 1 | 2 | 10 | ~4 | ~5 | ~6 | 20 [avatar split clause 950,2099,1563,1559,1555,1583,1547,1543]
2108. 8 | 10 | ~4 | ~5 | ~6 | 20 [avatar split clause 949,2099,1563,1559,1555,1583,1573]
2109. 9 | 10 | ~4 | ~5 | ~6 | 20 [avatar split clause 948,2099,1563,1559,1555,1583,1578]
2110. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | 20 [avatar split clause 947,2099,1563,1559,1555,1591,1547,1543]
2111. 1 | 2 | 11 | ~4 | ~5 | ~6 | 20 [avatar split clause 946,2099,1563,1559,1555,1591,1547,1543]
2112. 8 | 11 | ~4 | ~5 | ~6 | 20 [avatar split clause 945,2099,1563,1559,1555,1591,1573]
2113. 9 | 11 | ~4 | ~5 | ~6 | 20 [avatar split clause 944,2099,1563,1559,1555,1591,1578]
2114. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | 20 [avatar split clause 943,2099,1563,1559,1555,1599,1547,1543]
2115. 1 | 2 | 12 | ~4 | ~5 | ~6 | 20 [avatar split clause 942,2099,1563,1559,1555,1599,1547,1543]
2116. 8 | 12 | ~4 | ~5 | ~6 | 20 [avatar split clause 941,2099,1563,1559,1555,1599,1573]
2117. 9 | 12 | ~4 | ~5 | ~6 | 20 [avatar split clause 940,2099,1563,1559,1555,1599,1578]
2118. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | 20 [avatar split clause 939,2099,1563,1559,1607,1551,1547,1543]
2119. 1 | 2 | ~3 | 13 | ~5 | ~6 | 20 [avatar split clause 938,2099,1563,1559,1607,1551,1547,1543]
2120. 8 | ~3 | 13 | ~5 | ~6 | 20 [avatar split clause 937,2099,1563,1559,1607,1551,1573]
2121. 9 | ~3 | 13 | ~5 | ~6 | 20 [avatar split clause 936,2099,1563,1559,1607,1551,1578]
2122. ~1 | ~2 | 10 | 13 | ~5 | ~6 | 20 [avatar split clause 935,2099,1563,1559,1607,1583,1547,1543]
2123. 1 | 2 | 10 | 13 | ~5 | ~6 | 20 [avatar split clause 934,2099,1563,1559,1607,1583,1547,1543]
2124. 8 | 10 | 13 | ~5 | ~6 | 20 [avatar split clause 933,2099,1563,1559,1607,1583,1573]
2125. 9 | 10 | 13 | ~5 | ~6 | 20 [avatar split clause 932,2099,1563,1559,1607,1583,1578]
2126. ~1 | ~2 | 11 | 13 | ~5 | ~6 | 20 [avatar split clause 931,2099,1563,1559,1607,1591,1547,1543]
2127. 1 | 2 | 11 | 13 | ~5 | ~6 | 20 [avatar split clause 930,2099,1563,1559,1607,1591,1547,1543]
2128. 8 | 11 | 13 | ~5 | ~6 | 20 [avatar split clause 929,2099,1563,1559,1607,1591,1573]
2129. 9 | 11 | 13 | ~5 | ~6 | 20 [avatar split clause 928,2099,1563,1559,1607,1591,1578]
2130. ~1 | ~2 | 12 | 13 | ~5 | ~6 | 20 [avatar split clause 927,2099,1563,1559,1607,1599,1547,1543]
2131. 1 | 2 | 12 | 13 | ~5 | ~6 | 20 [avatar split clause 926,2099,1563,1559,1607,1599,1547,1543]
2132. 8 | 12 | 13 | ~5 | ~6 | 20 [avatar split clause 925,2099,1563,1559,1607,1599,1573]
2133. 9 | 12 | 13 | ~5 | ~6 | 20 [avatar split clause 924,2099,1563,1559,1607,1599,1578]
2134. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | 20 [avatar split clause 923,2099,1563,1559,1627,1551,1547,1543]
2135. 1 | 2 | ~3 | 14 | ~5 | ~6 | 20 [avatar split clause 922,2099,1563,1559,1627,1551,1547,1543]
2136. 8 | ~3 | 14 | ~5 | ~6 | 20 [avatar split clause 921,2099,1563,1559,1627,1551,1573]
2137. 9 | ~3 | 14 | ~5 | ~6 | 20 [avatar split clause 920,2099,1563,1559,1627,1551,1578]
2138. ~1 | ~2 | 10 | 14 | ~5 | ~6 | 20 [avatar split clause 919,2099,1563,1559,1627,1583,1547,1543]
2139. 1 | 2 | 10 | 14 | ~5 | ~6 | 20 [avatar split clause 918,2099,1563,1559,1627,1583,1547,1543]
2140. 8 | 10 | 14 | ~5 | ~6 | 20 [avatar split clause 917,2099,1563,1559,1627,1583,1573]
2141. 9 | 10 | 14 | ~5 | ~6 | 20 [avatar split clause 916,2099,1563,1559,1627,1583,1578]
2142. ~1 | ~2 | 11 | 14 | ~5 | ~6 | 20 [avatar split clause 915,2099,1563,1559,1627,1591,1547,1543]
2143. 1 | 2 | 11 | 14 | ~5 | ~6 | 20 [avatar split clause 914,2099,1563,1559,1627,1591,1547,1543]
2144. 8 | 11 | 14 | ~5 | ~6 | 20 [avatar split clause 913,2099,1563,1559,1627,1591,1573]
2145. 9 | 11 | 14 | ~5 | ~6 | 20 [avatar split clause 912,2099,1563,1559,1627,1591,1578]
2146. ~1 | ~2 | 12 | 14 | ~5 | ~6 | 20 [avatar split clause 911,2099,1563,1559,1627,1599,1547,1543]
2147. 1 | 2 | 12 | 14 | ~5 | ~6 | 20 [avatar split clause 910,2099,1563,1559,1627,1599,1547,1543]
2148. 8 | 12 | 14 | ~5 | ~6 | 20 [avatar split clause 909,2099,1563,1559,1627,1599,1573]
2149. 9 | 12 | 14 | ~5 | ~6 | 20 [avatar split clause 908,2099,1563,1559,1627,1599,1578]
2150. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | 20 [avatar split clause 907,2099,1563,1559,1647,1551,1547,1543]
2151. 1 | 2 | ~3 | 15 | ~5 | ~6 | 20 [avatar split clause 906,2099,1563,1559,1647,1551,1547,1543]
2152. 8 | ~3 | 15 | ~5 | ~6 | 20 [avatar split clause 905,2099,1563,1559,1647,1551,1573]
2153. 9 | ~3 | 15 | ~5 | ~6 | 20 [avatar split clause 904,2099,1563,1559,1647,1551,1578]
2154. ~1 | ~2 | 10 | 15 | ~5 | ~6 | 20 [avatar split clause 903,2099,1563,1559,1647,1583,1547,1543]
2155. 1 | 2 | 10 | 15 | ~5 | ~6 | 20 [avatar split clause 902,2099,1563,1559,1647,1583,1547,1543]
2156. 8 | 10 | 15 | ~5 | ~6 | 20 [avatar split clause 901,2099,1563,1559,1647,1583,1573]
2157. 9 | 10 | 15 | ~5 | ~6 | 20 [avatar split clause 900,2099,1563,1559,1647,1583,1578]
2158. ~1 | ~2 | 11 | 15 | ~5 | ~6 | 20 [avatar split clause 899,2099,1563,1559,1647,1591,1547,1543]
2159. 1 | 2 | 11 | 15 | ~5 | ~6 | 20 [avatar split clause 898,2099,1563,1559,1647,1591,1547,1543]
2160. 8 | 11 | 15 | ~5 | ~6 | 20 [avatar split clause 897,2099,1563,1559,1647,1591,1573]
2161. 9 | 11 | 15 | ~5 | ~6 | 20 [avatar split clause 896,2099,1563,1559,1647,1591,1578]
2162. ~1 | ~2 | 12 | 15 | ~5 | ~6 | 20 [avatar split clause 895,2099,1563,1559,1647,1599,1547,1543]
2163. 1 | 2 | 12 | 15 | ~5 | ~6 | 20 [avatar split clause 894,2099,1563,1559,1647,1599,1547,1543]
2164. 8 | 12 | 15 | ~5 | ~6 | 20 [avatar split clause 893,2099,1563,1559,1647,1599,1573]
2165. 9 | 12 | 15 | ~5 | ~6 | 20 [avatar split clause 892,2099,1563,1559,1647,1599,1578]
2166. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | 20 [avatar split clause 891,2099,1563,1559,1667,1551,1547,1543]
2167. 1 | 2 | ~3 | 16 | ~5 | ~6 | 20 [avatar split clause 890,2099,1563,1559,1667,1551,1547,1543]
2168. 8 | ~3 | 16 | ~5 | ~6 | 20 [avatar split clause 889,2099,1563,1559,1667,1551,1573]
2169. 9 | ~3 | 16 | ~5 | ~6 | 20 [avatar split clause 888,2099,1563,1559,1667,1551,1578]
2170. ~1 | ~2 | 10 | 16 | ~5 | ~6 | 20 [avatar split clause 887,2099,1563,1559,1667,1583,1547,1543]
2171. 1 | 2 | 10 | 16 | ~5 | ~6 | 20 [avatar split clause 886,2099,1563,1559,1667,1583,1547,1543]
2172. 8 | 10 | 16 | ~5 | ~6 | 20 [avatar split clause 885,2099,1563,1559,1667,1583,1573]
2173. 9 | 10 | 16 | ~5 | ~6 | 20 [avatar split clause 884,2099,1563,1559,1667,1583,1578]
2174. ~1 | ~2 | 11 | 16 | ~5 | ~6 | 20 [avatar split clause 883,2099,1563,1559,1667,1591,1547,1543]
2175. 1 | 2 | 11 | 16 | ~5 | ~6 | 20 [avatar split clause 882,2099,1563,1559,1667,1591,1547,1543]
2176. 8 | 11 | 16 | ~5 | ~6 | 20 [avatar split clause 881,2099,1563,1559,1667,1591,1573]
2177. 9 | 11 | 16 | ~5 | ~6 | 20 [avatar split clause 880,2099,1563,1559,1667,1591,1578]
2178. ~1 | ~2 | 12 | 16 | ~5 | ~6 | 20 [avatar split clause 879,2099,1563,1559,1667,1599,1547,1543]
2179. 1 | 2 | 12 | 16 | ~5 | ~6 | 20 [avatar split clause 878,2099,1563,1559,1667,1599,1547,1543]
2180. 8 | 12 | 16 | ~5 | ~6 | 20 [avatar split clause 877,2099,1563,1559,1667,1599,1573]
2181. 9 | 12 | 16 | ~5 | ~6 | 20 [avatar split clause 876,2099,1563,1559,1667,1599,1578]
2182. ~1 | ~2 | ~3 | ~4 | 17 | 20 [avatar split clause 875,2099,1687,1555,1551,1547,1543]
2183. 1 | 2 | ~3 | ~4 | 17 | 20 [avatar split clause 874,2099,1687,1555,1551,1547,1543]
2184. 8 | ~3 | ~4 | 17 | 20 [avatar split clause 873,2099,1687,1555,1551,1573]
2185. 9 | ~3 | ~4 | 17 | 20 [avatar split clause 872,2099,1687,1555,1551,1578]
2186. ~1 | ~2 | 10 | ~4 | 17 | 20 [avatar split clause 871,2099,1687,1555,1583,1547,1543]
2187. 1 | 2 | 10 | ~4 | 17 | 20 [avatar split clause 870,2099,1687,1555,1583,1547,1543]
2188. 8 | 10 | ~4 | 17 | 20 [avatar split clause 869,2099,1687,1555,1583,1573]
2189. 9 | 10 | ~4 | 17 | 20 [avatar split clause 868,2099,1687,1555,1583,1578]
2190. ~1 | ~2 | 11 | ~4 | 17 | 20 [avatar split clause 867,2099,1687,1555,1591,1547,1543]
2191. 1 | 2 | 11 | ~4 | 17 | 20 [avatar split clause 866,2099,1687,1555,1591,1547,1543]
2192. 8 | 11 | ~4 | 17 | 20 [avatar split clause 865,2099,1687,1555,1591,1573]
2193. 9 | 11 | ~4 | 17 | 20 [avatar split clause 864,2099,1687,1555,1591,1578]
2194. ~1 | ~2 | 12 | ~4 | 17 | 20 [avatar split clause 863,2099,1687,1555,1599,1547,1543]
2195. 1 | 2 | 12 | ~4 | 17 | 20 [avatar split clause 862,2099,1687,1555,1599,1547,1543]
2196. 8 | 12 | ~4 | 17 | 20 [avatar split clause 861,2099,1687,1555,1599,1573]
2197. 9 | 12 | ~4 | 17 | 20 [avatar split clause 860,2099,1687,1555,1599,1578]
2198. ~1 | ~2 | ~3 | 13 | 17 | 20 [avatar split clause 859,2099,1687,1607,1551,1547,1543]
2199. 1 | 2 | ~3 | 13 | 17 | 20 [avatar split clause 858,2099,1687,1607,1551,1547,1543]
2200. 8 | ~3 | 13 | 17 | 20 [avatar split clause 857,2099,1687,1607,1551,1573]
2201. 9 | ~3 | 13 | 17 | 20 [avatar split clause 856,2099,1687,1607,1551,1578]
2202. ~1 | ~2 | 10 | 13 | 17 | 20 [avatar split clause 855,2099,1687,1607,1583,1547,1543]
2203. 1 | 2 | 10 | 13 | 17 | 20 [avatar split clause 854,2099,1687,1607,1583,1547,1543]
2204. 8 | 10 | 13 | 17 | 20 [avatar split clause 853,2099,1687,1607,1583,1573]
2205. 9 | 10 | 13 | 17 | 20 [avatar split clause 852,2099,1687,1607,1583,1578]
2206. ~1 | ~2 | 11 | 13 | 17 | 20 [avatar split clause 851,2099,1687,1607,1591,1547,1543]
2207. 1 | 2 | 11 | 13 | 17 | 20 [avatar split clause 850,2099,1687,1607,1591,1547,1543]
2208. 8 | 11 | 13 | 17 | 20 [avatar split clause 849,2099,1687,1607,1591,1573]
2209. 9 | 11 | 13 | 17 | 20 [avatar split clause 848,2099,1687,1607,1591,1578]
2210. ~1 | ~2 | 12 | 13 | 17 | 20 [avatar split clause 847,2099,1687,1607,1599,1547,1543]
2211. 1 | 2 | 12 | 13 | 17 | 20 [avatar split clause 846,2099,1687,1607,1599,1547,1543]
2212. 8 | 12 | 13 | 17 | 20 [avatar split clause 845,2099,1687,1607,1599,1573]
2213. 9 | 12 | 13 | 17 | 20 [avatar split clause 844,2099,1687,1607,1599,1578]
2214. ~1 | ~2 | ~3 | 14 | 17 | 20 [avatar split clause 843,2099,1687,1627,1551,1547,1543]
2215. 1 | 2 | ~3 | 14 | 17 | 20 [avatar split clause 842,2099,1687,1627,1551,1547,1543]
2216. 8 | ~3 | 14 | 17 | 20 [avatar split clause 841,2099,1687,1627,1551,1573]
2217. 9 | ~3 | 14 | 17 | 20 [avatar split clause 840,2099,1687,1627,1551,1578]
2218. ~1 | ~2 | 10 | 14 | 17 | 20 [avatar split clause 839,2099,1687,1627,1583,1547,1543]
2219. 1 | 2 | 10 | 14 | 17 | 20 [avatar split clause 838,2099,1687,1627,1583,1547,1543]
2220. 8 | 10 | 14 | 17 | 20 [avatar split clause 837,2099,1687,1627,1583,1573]
2221. 9 | 10 | 14 | 17 | 20 [avatar split clause 836,2099,1687,1627,1583,1578]
2222. ~1 | ~2 | 11 | 14 | 17 | 20 [avatar split clause 835,2099,1687,1627,1591,1547,1543]
2223. 1 | 2 | 11 | 14 | 17 | 20 [avatar split clause 834,2099,1687,1627,1591,1547,1543]
2224. 8 | 11 | 14 | 17 | 20 [avatar split clause 833,2099,1687,1627,1591,1573]
2225. 9 | 11 | 14 | 17 | 20 [avatar split clause 832,2099,1687,1627,1591,1578]
2226. ~1 | ~2 | 12 | 14 | 17 | 20 [avatar split clause 831,2099,1687,1627,1599,1547,1543]
2227. 1 | 2 | 12 | 14 | 17 | 20 [avatar split clause 830,2099,1687,1627,1599,1547,1543]
2228. 8 | 12 | 14 | 17 | 20 [avatar split clause 829,2099,1687,1627,1599,1573]
2229. 9 | 12 | 14 | 17 | 20 [avatar split clause 828,2099,1687,1627,1599,1578]
2230. ~1 | ~2 | ~3 | 15 | 17 | 20 [avatar split clause 827,2099,1687,1647,1551,1547,1543]
2231. 1 | 2 | ~3 | 15 | 17 | 20 [avatar split clause 826,2099,1687,1647,1551,1547,1543]
2232. 8 | ~3 | 15 | 17 | 20 [avatar split clause 825,2099,1687,1647,1551,1573]
2233. 9 | ~3 | 15 | 17 | 20 [avatar split clause 824,2099,1687,1647,1551,1578]
2234. ~1 | ~2 | 10 | 15 | 17 | 20 [avatar split clause 823,2099,1687,1647,1583,1547,1543]
2235. 1 | 2 | 10 | 15 | 17 | 20 [avatar split clause 822,2099,1687,1647,1583,1547,1543]
2236. 8 | 10 | 15 | 17 | 20 [avatar split clause 821,2099,1687,1647,1583,1573]
2237. 9 | 10 | 15 | 17 | 20 [avatar split clause 820,2099,1687,1647,1583,1578]
2238. ~1 | ~2 | 11 | 15 | 17 | 20 [avatar split clause 819,2099,1687,1647,1591,1547,1543]
2239. 1 | 2 | 11 | 15 | 17 | 20 [avatar split clause 818,2099,1687,1647,1591,1547,1543]
2240. 8 | 11 | 15 | 17 | 20 [avatar split clause 817,2099,1687,1647,1591,1573]
2241. 9 | 11 | 15 | 17 | 20 [avatar split clause 816,2099,1687,1647,1591,1578]
2242. ~1 | ~2 | 12 | 15 | 17 | 20 [avatar split clause 815,2099,1687,1647,1599,1547,1543]
2243. 1 | 2 | 12 | 15 | 17 | 20 [avatar split clause 814,2099,1687,1647,1599,1547,1543]
2244. 8 | 12 | 15 | 17 | 20 [avatar split clause 813,2099,1687,1647,1599,1573]
2245. 9 | 12 | 15 | 17 | 20 [avatar split clause 812,2099,1687,1647,1599,1578]
2246. ~1 | ~2 | ~3 | 16 | 17 | 20 [avatar split clause 811,2099,1687,1667,1551,1547,1543]
2247. 1 | 2 | ~3 | 16 | 17 | 20 [avatar split clause 810,2099,1687,1667,1551,1547,1543]
2248. 8 | ~3 | 16 | 17 | 20 [avatar split clause 809,2099,1687,1667,1551,1573]
2249. 9 | ~3 | 16 | 17 | 20 [avatar split clause 808,2099,1687,1667,1551,1578]
2250. ~1 | ~2 | 10 | 16 | 17 | 20 [avatar split clause 807,2099,1687,1667,1583,1547,1543]
2251. 1 | 2 | 10 | 16 | 17 | 20 [avatar split clause 806,2099,1687,1667,1583,1547,1543]
2252. 8 | 10 | 16 | 17 | 20 [avatar split clause 805,2099,1687,1667,1583,1573]
2253. 9 | 10 | 16 | 17 | 20 [avatar split clause 804,2099,1687,1667,1583,1578]
2254. ~1 | ~2 | 11 | 16 | 17 | 20 [avatar split clause 803,2099,1687,1667,1591,1547,1543]
2255. 1 | 2 | 11 | 16 | 17 | 20 [avatar split clause 802,2099,1687,1667,1591,1547,1543]
2256. 8 | 11 | 16 | 17 | 20 [avatar split clause 801,2099,1687,1667,1591,1573]
2257. 9 | 11 | 16 | 17 | 20 [avatar split clause 800,2099,1687,1667,1591,1578]
2258. ~1 | ~2 | 12 | 16 | 17 | 20 [avatar split clause 799,2099,1687,1667,1599,1547,1543]
2259. 1 | 2 | 12 | 16 | 17 | 20 [avatar split clause 798,2099,1687,1667,1599,1547,1543]
2260. 8 | 12 | 16 | 17 | 20 [avatar split clause 797,2099,1687,1667,1599,1573]
2261. 9 | 12 | 16 | 17 | 20 [avatar split clause 796,2099,1687,1667,1599,1578]
2262. ~1 | ~2 | ~3 | ~4 | 18 | 20 [avatar split clause 795,2099,1771,1555,1551,1547,1543]
2263. 1 | 2 | ~3 | ~4 | 18 | 20 [avatar split clause 794,2099,1771,1555,1551,1547,1543]
2264. 8 | ~3 | ~4 | 18 | 20 [avatar split clause 793,2099,1771,1555,1551,1573]
2265. 9 | ~3 | ~4 | 18 | 20 [avatar split clause 792,2099,1771,1555,1551,1578]
2266. ~1 | ~2 | 10 | ~4 | 18 | 20 [avatar split clause 791,2099,1771,1555,1583,1547,1543]
2267. 1 | 2 | 10 | ~4 | 18 | 20 [avatar split clause 790,2099,1771,1555,1583,1547,1543]
2268. 8 | 10 | ~4 | 18 | 20 [avatar split clause 789,2099,1771,1555,1583,1573]
2269. 9 | 10 | ~4 | 18 | 20 [avatar split clause 788,2099,1771,1555,1583,1578]
2270. ~1 | ~2 | 11 | ~4 | 18 | 20 [avatar split clause 787,2099,1771,1555,1591,1547,1543]
2271. 1 | 2 | 11 | ~4 | 18 | 20 [avatar split clause 786,2099,1771,1555,1591,1547,1543]
2272. 8 | 11 | ~4 | 18 | 20 [avatar split clause 785,2099,1771,1555,1591,1573]
2273. 9 | 11 | ~4 | 18 | 20 [avatar split clause 784,2099,1771,1555,1591,1578]
2274. ~1 | ~2 | 12 | ~4 | 18 | 20 [avatar split clause 783,2099,1771,1555,1599,1547,1543]
2275. 1 | 2 | 12 | ~4 | 18 | 20 [avatar split clause 782,2099,1771,1555,1599,1547,1543]
2276. 8 | 12 | ~4 | 18 | 20 [avatar split clause 781,2099,1771,1555,1599,1573]
2277. 9 | 12 | ~4 | 18 | 20 [avatar split clause 780,2099,1771,1555,1599,1578]
2278. ~1 | ~2 | ~3 | 13 | 18 | 20 [avatar split clause 779,2099,1771,1607,1551,1547,1543]
2279. 1 | 2 | ~3 | 13 | 18 | 20 [avatar split clause 778,2099,1771,1607,1551,1547,1543]
2280. 8 | ~3 | 13 | 18 | 20 [avatar split clause 777,2099,1771,1607,1551,1573]
2281. 9 | ~3 | 13 | 18 | 20 [avatar split clause 776,2099,1771,1607,1551,1578]
2282. ~1 | ~2 | 10 | 13 | 18 | 20 [avatar split clause 775,2099,1771,1607,1583,1547,1543]
2283. 1 | 2 | 10 | 13 | 18 | 20 [avatar split clause 774,2099,1771,1607,1583,1547,1543]
2284. 8 | 10 | 13 | 18 | 20 [avatar split clause 773,2099,1771,1607,1583,1573]
2285. 9 | 10 | 13 | 18 | 20 [avatar split clause 772,2099,1771,1607,1583,1578]
2286. ~1 | ~2 | 11 | 13 | 18 | 20 [avatar split clause 771,2099,1771,1607,1591,1547,1543]
2287. 1 | 2 | 11 | 13 | 18 | 20 [avatar split clause 770,2099,1771,1607,1591,1547,1543]
2288. 8 | 11 | 13 | 18 | 20 [avatar split clause 769,2099,1771,1607,1591,1573]
2289. 9 | 11 | 13 | 18 | 20 [avatar split clause 768,2099,1771,1607,1591,1578]
2290. ~1 | ~2 | 12 | 13 | 18 | 20 [avatar split clause 767,2099,1771,1607,1599,1547,1543]
2291. 1 | 2 | 12 | 13 | 18 | 20 [avatar split clause 766,2099,1771,1607,1599,1547,1543]
2292. 8 | 12 | 13 | 18 | 20 [avatar split clause 765,2099,1771,1607,1599,1573]
2293. 9 | 12 | 13 | 18 | 20 [avatar split clause 764,2099,1771,1607,1599,1578]
2294. ~1 | ~2 | ~3 | 14 | 18 | 20 [avatar split clause 763,2099,1771,1627,1551,1547,1543]
2295. 1 | 2 | ~3 | 14 | 18 | 20 [avatar split clause 762,2099,1771,1627,1551,1547,1543]
2296. 8 | ~3 | 14 | 18 | 20 [avatar split clause 761,2099,1771,1627,1551,1573]
2297. 9 | ~3 | 14 | 18 | 20 [avatar split clause 760,2099,1771,1627,1551,1578]
2298. ~1 | ~2 | 10 | 14 | 18 | 20 [avatar split clause 759,2099,1771,1627,1583,1547,1543]
2299. 1 | 2 | 10 | 14 | 18 | 20 [avatar split clause 758,2099,1771,1627,1583,1547,1543]
2300. 8 | 10 | 14 | 18 | 20 [avatar split clause 757,2099,1771,1627,1583,1573]
2301. 9 | 10 | 14 | 18 | 20 [avatar split clause 756,2099,1771,1627,1583,1578]
2302. ~1 | ~2 | 11 | 14 | 18 | 20 [avatar split clause 755,2099,1771,1627,1591,1547,1543]
2303. 1 | 2 | 11 | 14 | 18 | 20 [avatar split clause 754,2099,1771,1627,1591,1547,1543]
2304. 8 | 11 | 14 | 18 | 20 [avatar split clause 753,2099,1771,1627,1591,1573]
2305. 9 | 11 | 14 | 18 | 20 [avatar split clause 752,2099,1771,1627,1591,1578]
2306. ~1 | ~2 | 12 | 14 | 18 | 20 [avatar split clause 751,2099,1771,1627,1599,1547,1543]
2307. 1 | 2 | 12 | 14 | 18 | 20 [avatar split clause 750,2099,1771,1627,1599,1547,1543]
2308. 8 | 12 | 14 | 18 | 20 [avatar split clause 749,2099,1771,1627,1599,1573]
2309. 9 | 12 | 14 | 18 | 20 [avatar split clause 748,2099,1771,1627,1599,1578]
2310. ~1 | ~2 | ~3 | 15 | 18 | 20 [avatar split clause 747,2099,1771,1647,1551,1547,1543]
2311. 1 | 2 | ~3 | 15 | 18 | 20 [avatar split clause 746,2099,1771,1647,1551,1547,1543]
2312. 8 | ~3 | 15 | 18 | 20 [avatar split clause 745,2099,1771,1647,1551,1573]
2313. 9 | ~3 | 15 | 18 | 20 [avatar split clause 744,2099,1771,1647,1551,1578]
2314. ~1 | ~2 | 10 | 15 | 18 | 20 [avatar split clause 743,2099,1771,1647,1583,1547,1543]
2315. 1 | 2 | 10 | 15 | 18 | 20 [avatar split clause 742,2099,1771,1647,1583,1547,1543]
2316. 8 | 10 | 15 | 18 | 20 [avatar split clause 741,2099,1771,1647,1583,1573]
2317. 9 | 10 | 15 | 18 | 20 [avatar split clause 740,2099,1771,1647,1583,1578]
2318. ~1 | ~2 | 11 | 15 | 18 | 20 [avatar split clause 739,2099,1771,1647,1591,1547,1543]
2319. 1 | 2 | 11 | 15 | 18 | 20 [avatar split clause 738,2099,1771,1647,1591,1547,1543]
2320. 8 | 11 | 15 | 18 | 20 [avatar split clause 737,2099,1771,1647,1591,1573]
2321. 9 | 11 | 15 | 18 | 20 [avatar split clause 736,2099,1771,1647,1591,1578]
2322. ~1 | ~2 | 12 | 15 | 18 | 20 [avatar split clause 735,2099,1771,1647,1599,1547,1543]
2323. 1 | 2 | 12 | 15 | 18 | 20 [avatar split clause 734,2099,1771,1647,1599,1547,1543]
2324. 8 | 12 | 15 | 18 | 20 [avatar split clause 733,2099,1771,1647,1599,1573]
2325. 9 | 12 | 15 | 18 | 20 [avatar split clause 732,2099,1771,1647,1599,1578]
2326. ~1 | ~2 | ~3 | 16 | 18 | 20 [avatar split clause 731,2099,1771,1667,1551,1547,1543]
2327. 1 | 2 | ~3 | 16 | 18 | 20 [avatar split clause 730,2099,1771,1667,1551,1547,1543]
2328. 8 | ~3 | 16 | 18 | 20 [avatar split clause 729,2099,1771,1667,1551,1573]
2329. 9 | ~3 | 16 | 18 | 20 [avatar split clause 728,2099,1771,1667,1551,1578]
2330. ~1 | ~2 | 10 | 16 | 18 | 20 [avatar split clause 727,2099,1771,1667,1583,1547,1543]
2331. 1 | 2 | 10 | 16 | 18 | 20 [avatar split clause 726,2099,1771,1667,1583,1547,1543]
2332. 8 | 10 | 16 | 18 | 20 [avatar split clause 725,2099,1771,1667,1583,1573]
2333. 9 | 10 | 16 | 18 | 20 [avatar split clause 724,2099,1771,1667,1583,1578]
2334. ~1 | ~2 | 11 | 16 | 18 | 20 [avatar split clause 723,2099,1771,1667,1591,1547,1543]
2335. 1 | 2 | 11 | 16 | 18 | 20 [avatar split clause 722,2099,1771,1667,1591,1547,1543]
2336. 8 | 11 | 16 | 18 | 20 [avatar split clause 721,2099,1771,1667,1591,1573]
2337. 9 | 11 | 16 | 18 | 20 [avatar split clause 720,2099,1771,1667,1591,1578]
2338. ~1 | ~2 | 12 | 16 | 18 | 20 [avatar split clause 719,2099,1771,1667,1599,1547,1543]
2339. 1 | 2 | 12 | 16 | 18 | 20 [avatar split clause 718,2099,1771,1667,1599,1547,1543]
2340. 8 | 12 | 16 | 18 | 20 [avatar split clause 717,2099,1771,1667,1599,1573]
2341. 9 | 12 | 16 | 18 | 20 [avatar split clause 716,2099,1771,1667,1599,1578]
2343. 21 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
2346. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | 21 [avatar split clause 715,2343,1563,1559,1555,1551,1547,1543]
2347. 1 | 2 | ~3 | ~4 | ~5 | ~6 | 21 [avatar split clause 714,2343,1563,1559,1555,1551,1547,1543]
2348. 8 | ~3 | ~4 | ~5 | ~6 | 21 [avatar split clause 713,2343,1563,1559,1555,1551,1573]
2349. 9 | ~3 | ~4 | ~5 | ~6 | 21 [avatar split clause 712,2343,1563,1559,1555,1551,1578]
2350. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | 21 [avatar split clause 711,2343,1563,1559,1555,1583,1547,1543]
2351. 1 | 2 | 10 | ~4 | ~5 | ~6 | 21 [avatar split clause 710,2343,1563,1559,1555,1583,1547,1543]
2352. 8 | 10 | ~4 | ~5 | ~6 | 21 [avatar split clause 709,2343,1563,1559,1555,1583,1573]
2353. 9 | 10 | ~4 | ~5 | ~6 | 21 [avatar split clause 708,2343,1563,1559,1555,1583,1578]
2354. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | 21 [avatar split clause 707,2343,1563,1559,1555,1591,1547,1543]
2355. 1 | 2 | 11 | ~4 | ~5 | ~6 | 21 [avatar split clause 706,2343,1563,1559,1555,1591,1547,1543]
2356. 8 | 11 | ~4 | ~5 | ~6 | 21 [avatar split clause 705,2343,1563,1559,1555,1591,1573]
2357. 9 | 11 | ~4 | ~5 | ~6 | 21 [avatar split clause 704,2343,1563,1559,1555,1591,1578]
2358. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | 21 [avatar split clause 703,2343,1563,1559,1555,1599,1547,1543]
2359. 1 | 2 | 12 | ~4 | ~5 | ~6 | 21 [avatar split clause 702,2343,1563,1559,1555,1599,1547,1543]
2360. 8 | 12 | ~4 | ~5 | ~6 | 21 [avatar split clause 701,2343,1563,1559,1555,1599,1573]
2361. 9 | 12 | ~4 | ~5 | ~6 | 21 [avatar split clause 700,2343,1563,1559,1555,1599,1578]
2362. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | 21 [avatar split clause 699,2343,1563,1559,1607,1551,1547,1543]
2363. 1 | 2 | ~3 | 13 | ~5 | ~6 | 21 [avatar split clause 698,2343,1563,1559,1607,1551,1547,1543]
2364. 8 | ~3 | 13 | ~5 | ~6 | 21 [avatar split clause 697,2343,1563,1559,1607,1551,1573]
2365. 9 | ~3 | 13 | ~5 | ~6 | 21 [avatar split clause 696,2343,1563,1559,1607,1551,1578]
2366. ~1 | ~2 | 10 | 13 | ~5 | ~6 | 21 [avatar split clause 695,2343,1563,1559,1607,1583,1547,1543]
2367. 1 | 2 | 10 | 13 | ~5 | ~6 | 21 [avatar split clause 694,2343,1563,1559,1607,1583,1547,1543]
2368. 8 | 10 | 13 | ~5 | ~6 | 21 [avatar split clause 693,2343,1563,1559,1607,1583,1573]
2369. 9 | 10 | 13 | ~5 | ~6 | 21 [avatar split clause 692,2343,1563,1559,1607,1583,1578]
2370. ~1 | ~2 | 11 | 13 | ~5 | ~6 | 21 [avatar split clause 691,2343,1563,1559,1607,1591,1547,1543]
2371. 1 | 2 | 11 | 13 | ~5 | ~6 | 21 [avatar split clause 690,2343,1563,1559,1607,1591,1547,1543]
2372. 8 | 11 | 13 | ~5 | ~6 | 21 [avatar split clause 689,2343,1563,1559,1607,1591,1573]
2373. 9 | 11 | 13 | ~5 | ~6 | 21 [avatar split clause 688,2343,1563,1559,1607,1591,1578]
2374. ~1 | ~2 | 12 | 13 | ~5 | ~6 | 21 [avatar split clause 687,2343,1563,1559,1607,1599,1547,1543]
2375. 1 | 2 | 12 | 13 | ~5 | ~6 | 21 [avatar split clause 686,2343,1563,1559,1607,1599,1547,1543]
2376. 8 | 12 | 13 | ~5 | ~6 | 21 [avatar split clause 685,2343,1563,1559,1607,1599,1573]
2377. 9 | 12 | 13 | ~5 | ~6 | 21 [avatar split clause 684,2343,1563,1559,1607,1599,1578]
2378. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | 21 [avatar split clause 683,2343,1563,1559,1627,1551,1547,1543]
2379. 1 | 2 | ~3 | 14 | ~5 | ~6 | 21 [avatar split clause 682,2343,1563,1559,1627,1551,1547,1543]
2380. 8 | ~3 | 14 | ~5 | ~6 | 21 [avatar split clause 681,2343,1563,1559,1627,1551,1573]
2381. 9 | ~3 | 14 | ~5 | ~6 | 21 [avatar split clause 680,2343,1563,1559,1627,1551,1578]
2382. ~1 | ~2 | 10 | 14 | ~5 | ~6 | 21 [avatar split clause 679,2343,1563,1559,1627,1583,1547,1543]
2383. 1 | 2 | 10 | 14 | ~5 | ~6 | 21 [avatar split clause 678,2343,1563,1559,1627,1583,1547,1543]
2384. 8 | 10 | 14 | ~5 | ~6 | 21 [avatar split clause 677,2343,1563,1559,1627,1583,1573]
2385. 9 | 10 | 14 | ~5 | ~6 | 21 [avatar split clause 676,2343,1563,1559,1627,1583,1578]
2386. ~1 | ~2 | 11 | 14 | ~5 | ~6 | 21 [avatar split clause 675,2343,1563,1559,1627,1591,1547,1543]
2387. 1 | 2 | 11 | 14 | ~5 | ~6 | 21 [avatar split clause 674,2343,1563,1559,1627,1591,1547,1543]
2388. 8 | 11 | 14 | ~5 | ~6 | 21 [avatar split clause 673,2343,1563,1559,1627,1591,1573]
2389. 9 | 11 | 14 | ~5 | ~6 | 21 [avatar split clause 672,2343,1563,1559,1627,1591,1578]
2390. ~1 | ~2 | 12 | 14 | ~5 | ~6 | 21 [avatar split clause 671,2343,1563,1559,1627,1599,1547,1543]
2391. 1 | 2 | 12 | 14 | ~5 | ~6 | 21 [avatar split clause 670,2343,1563,1559,1627,1599,1547,1543]
2392. 8 | 12 | 14 | ~5 | ~6 | 21 [avatar split clause 669,2343,1563,1559,1627,1599,1573]
2393. 9 | 12 | 14 | ~5 | ~6 | 21 [avatar split clause 668,2343,1563,1559,1627,1599,1578]
2394. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | 21 [avatar split clause 667,2343,1563,1559,1647,1551,1547,1543]
2395. 1 | 2 | ~3 | 15 | ~5 | ~6 | 21 [avatar split clause 666,2343,1563,1559,1647,1551,1547,1543]
2396. 8 | ~3 | 15 | ~5 | ~6 | 21 [avatar split clause 665,2343,1563,1559,1647,1551,1573]
2397. 9 | ~3 | 15 | ~5 | ~6 | 21 [avatar split clause 664,2343,1563,1559,1647,1551,1578]
2398. ~1 | ~2 | 10 | 15 | ~5 | ~6 | 21 [avatar split clause 663,2343,1563,1559,1647,1583,1547,1543]
2399. 1 | 2 | 10 | 15 | ~5 | ~6 | 21 [avatar split clause 662,2343,1563,1559,1647,1583,1547,1543]
2400. 8 | 10 | 15 | ~5 | ~6 | 21 [avatar split clause 661,2343,1563,1559,1647,1583,1573]
2401. 9 | 10 | 15 | ~5 | ~6 | 21 [avatar split clause 660,2343,1563,1559,1647,1583,1578]
2402. ~1 | ~2 | 11 | 15 | ~5 | ~6 | 21 [avatar split clause 659,2343,1563,1559,1647,1591,1547,1543]
2403. 1 | 2 | 11 | 15 | ~5 | ~6 | 21 [avatar split clause 658,2343,1563,1559,1647,1591,1547,1543]
2404. 8 | 11 | 15 | ~5 | ~6 | 21 [avatar split clause 657,2343,1563,1559,1647,1591,1573]
2405. 9 | 11 | 15 | ~5 | ~6 | 21 [avatar split clause 656,2343,1563,1559,1647,1591,1578]
2406. ~1 | ~2 | 12 | 15 | ~5 | ~6 | 21 [avatar split clause 655,2343,1563,1559,1647,1599,1547,1543]
2407. 1 | 2 | 12 | 15 | ~5 | ~6 | 21 [avatar split clause 654,2343,1563,1559,1647,1599,1547,1543]
2408. 8 | 12 | 15 | ~5 | ~6 | 21 [avatar split clause 653,2343,1563,1559,1647,1599,1573]
2409. 9 | 12 | 15 | ~5 | ~6 | 21 [avatar split clause 652,2343,1563,1559,1647,1599,1578]
2410. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | 21 [avatar split clause 651,2343,1563,1559,1667,1551,1547,1543]
2411. 1 | 2 | ~3 | 16 | ~5 | ~6 | 21 [avatar split clause 650,2343,1563,1559,1667,1551,1547,1543]
2412. 8 | ~3 | 16 | ~5 | ~6 | 21 [avatar split clause 649,2343,1563,1559,1667,1551,1573]
2413. 9 | ~3 | 16 | ~5 | ~6 | 21 [avatar split clause 648,2343,1563,1559,1667,1551,1578]
2414. ~1 | ~2 | 10 | 16 | ~5 | ~6 | 21 [avatar split clause 647,2343,1563,1559,1667,1583,1547,1543]
2415. 1 | 2 | 10 | 16 | ~5 | ~6 | 21 [avatar split clause 646,2343,1563,1559,1667,1583,1547,1543]
2416. 8 | 10 | 16 | ~5 | ~6 | 21 [avatar split clause 645,2343,1563,1559,1667,1583,1573]
2417. 9 | 10 | 16 | ~5 | ~6 | 21 [avatar split clause 644,2343,1563,1559,1667,1583,1578]
2418. ~1 | ~2 | 11 | 16 | ~5 | ~6 | 21 [avatar split clause 643,2343,1563,1559,1667,1591,1547,1543]
2419. 1 | 2 | 11 | 16 | ~5 | ~6 | 21 [avatar split clause 642,2343,1563,1559,1667,1591,1547,1543]
2420. 8 | 11 | 16 | ~5 | ~6 | 21 [avatar split clause 641,2343,1563,1559,1667,1591,1573]
2421. 9 | 11 | 16 | ~5 | ~6 | 21 [avatar split clause 640,2343,1563,1559,1667,1591,1578]
2422. ~1 | ~2 | 12 | 16 | ~5 | ~6 | 21 [avatar split clause 639,2343,1563,1559,1667,1599,1547,1543]
2423. 1 | 2 | 12 | 16 | ~5 | ~6 | 21 [avatar split clause 638,2343,1563,1559,1667,1599,1547,1543]
2424. 8 | 12 | 16 | ~5 | ~6 | 21 [avatar split clause 637,2343,1563,1559,1667,1599,1573]
2425. 9 | 12 | 16 | ~5 | ~6 | 21 [avatar split clause 636,2343,1563,1559,1667,1599,1578]
2426. ~1 | ~2 | ~3 | ~4 | 17 | 21 [avatar split clause 635,2343,1687,1555,1551,1547,1543]
2427. 1 | 2 | ~3 | ~4 | 17 | 21 [avatar split clause 634,2343,1687,1555,1551,1547,1543]
2428. 8 | ~3 | ~4 | 17 | 21 [avatar split clause 633,2343,1687,1555,1551,1573]
2429. 9 | ~3 | ~4 | 17 | 21 [avatar split clause 632,2343,1687,1555,1551,1578]
2430. ~1 | ~2 | 10 | ~4 | 17 | 21 [avatar split clause 631,2343,1687,1555,1583,1547,1543]
2431. 1 | 2 | 10 | ~4 | 17 | 21 [avatar split clause 630,2343,1687,1555,1583,1547,1543]
2432. 8 | 10 | ~4 | 17 | 21 [avatar split clause 629,2343,1687,1555,1583,1573]
2433. 9 | 10 | ~4 | 17 | 21 [avatar split clause 628,2343,1687,1555,1583,1578]
2434. ~1 | ~2 | 11 | ~4 | 17 | 21 [avatar split clause 627,2343,1687,1555,1591,1547,1543]
2435. 1 | 2 | 11 | ~4 | 17 | 21 [avatar split clause 626,2343,1687,1555,1591,1547,1543]
2436. 8 | 11 | ~4 | 17 | 21 [avatar split clause 625,2343,1687,1555,1591,1573]
2437. 9 | 11 | ~4 | 17 | 21 [avatar split clause 624,2343,1687,1555,1591,1578]
2438. ~1 | ~2 | 12 | ~4 | 17 | 21 [avatar split clause 623,2343,1687,1555,1599,1547,1543]
2439. 1 | 2 | 12 | ~4 | 17 | 21 [avatar split clause 622,2343,1687,1555,1599,1547,1543]
2440. 8 | 12 | ~4 | 17 | 21 [avatar split clause 621,2343,1687,1555,1599,1573]
2441. 9 | 12 | ~4 | 17 | 21 [avatar split clause 620,2343,1687,1555,1599,1578]
2442. ~1 | ~2 | ~3 | 13 | 17 | 21 [avatar split clause 619,2343,1687,1607,1551,1547,1543]
2443. 1 | 2 | ~3 | 13 | 17 | 21 [avatar split clause 618,2343,1687,1607,1551,1547,1543]
2444. 8 | ~3 | 13 | 17 | 21 [avatar split clause 617,2343,1687,1607,1551,1573]
2445. 9 | ~3 | 13 | 17 | 21 [avatar split clause 616,2343,1687,1607,1551,1578]
2446. ~1 | ~2 | 10 | 13 | 17 | 21 [avatar split clause 615,2343,1687,1607,1583,1547,1543]
2447. 1 | 2 | 10 | 13 | 17 | 21 [avatar split clause 614,2343,1687,1607,1583,1547,1543]
2448. 8 | 10 | 13 | 17 | 21 [avatar split clause 613,2343,1687,1607,1583,1573]
2449. 9 | 10 | 13 | 17 | 21 [avatar split clause 612,2343,1687,1607,1583,1578]
2450. ~1 | ~2 | 11 | 13 | 17 | 21 [avatar split clause 611,2343,1687,1607,1591,1547,1543]
2451. 1 | 2 | 11 | 13 | 17 | 21 [avatar split clause 610,2343,1687,1607,1591,1547,1543]
2452. 8 | 11 | 13 | 17 | 21 [avatar split clause 609,2343,1687,1607,1591,1573]
2453. 9 | 11 | 13 | 17 | 21 [avatar split clause 608,2343,1687,1607,1591,1578]
2454. ~1 | ~2 | 12 | 13 | 17 | 21 [avatar split clause 607,2343,1687,1607,1599,1547,1543]
2455. 1 | 2 | 12 | 13 | 17 | 21 [avatar split clause 606,2343,1687,1607,1599,1547,1543]
2456. 8 | 12 | 13 | 17 | 21 [avatar split clause 605,2343,1687,1607,1599,1573]
2457. 9 | 12 | 13 | 17 | 21 [avatar split clause 604,2343,1687,1607,1599,1578]
2458. ~1 | ~2 | ~3 | 14 | 17 | 21 [avatar split clause 603,2343,1687,1627,1551,1547,1543]
2459. 1 | 2 | ~3 | 14 | 17 | 21 [avatar split clause 602,2343,1687,1627,1551,1547,1543]
2460. 8 | ~3 | 14 | 17 | 21 [avatar split clause 601,2343,1687,1627,1551,1573]
2461. 9 | ~3 | 14 | 17 | 21 [avatar split clause 600,2343,1687,1627,1551,1578]
2462. ~1 | ~2 | 10 | 14 | 17 | 21 [avatar split clause 599,2343,1687,1627,1583,1547,1543]
2463. 1 | 2 | 10 | 14 | 17 | 21 [avatar split clause 598,2343,1687,1627,1583,1547,1543]
2464. 8 | 10 | 14 | 17 | 21 [avatar split clause 597,2343,1687,1627,1583,1573]
2465. 9 | 10 | 14 | 17 | 21 [avatar split clause 596,2343,1687,1627,1583,1578]
2466. ~1 | ~2 | 11 | 14 | 17 | 21 [avatar split clause 595,2343,1687,1627,1591,1547,1543]
2467. 1 | 2 | 11 | 14 | 17 | 21 [avatar split clause 594,2343,1687,1627,1591,1547,1543]
2468. 8 | 11 | 14 | 17 | 21 [avatar split clause 593,2343,1687,1627,1591,1573]
2469. 9 | 11 | 14 | 17 | 21 [avatar split clause 592,2343,1687,1627,1591,1578]
2470. ~1 | ~2 | 12 | 14 | 17 | 21 [avatar split clause 591,2343,1687,1627,1599,1547,1543]
2471. 1 | 2 | 12 | 14 | 17 | 21 [avatar split clause 590,2343,1687,1627,1599,1547,1543]
2472. 8 | 12 | 14 | 17 | 21 [avatar split clause 589,2343,1687,1627,1599,1573]
2473. 9 | 12 | 14 | 17 | 21 [avatar split clause 588,2343,1687,1627,1599,1578]
2474. ~1 | ~2 | ~3 | 15 | 17 | 21 [avatar split clause 587,2343,1687,1647,1551,1547,1543]
2475. 1 | 2 | ~3 | 15 | 17 | 21 [avatar split clause 586,2343,1687,1647,1551,1547,1543]
2476. 8 | ~3 | 15 | 17 | 21 [avatar split clause 585,2343,1687,1647,1551,1573]
2477. 9 | ~3 | 15 | 17 | 21 [avatar split clause 584,2343,1687,1647,1551,1578]
2478. ~1 | ~2 | 10 | 15 | 17 | 21 [avatar split clause 583,2343,1687,1647,1583,1547,1543]
2479. 1 | 2 | 10 | 15 | 17 | 21 [avatar split clause 582,2343,1687,1647,1583,1547,1543]
2480. 8 | 10 | 15 | 17 | 21 [avatar split clause 581,2343,1687,1647,1583,1573]
2481. 9 | 10 | 15 | 17 | 21 [avatar split clause 580,2343,1687,1647,1583,1578]
2482. ~1 | ~2 | 11 | 15 | 17 | 21 [avatar split clause 579,2343,1687,1647,1591,1547,1543]
2483. 1 | 2 | 11 | 15 | 17 | 21 [avatar split clause 578,2343,1687,1647,1591,1547,1543]
2484. 8 | 11 | 15 | 17 | 21 [avatar split clause 577,2343,1687,1647,1591,1573]
2485. 9 | 11 | 15 | 17 | 21 [avatar split clause 576,2343,1687,1647,1591,1578]
2486. ~1 | ~2 | 12 | 15 | 17 | 21 [avatar split clause 575,2343,1687,1647,1599,1547,1543]
2487. 1 | 2 | 12 | 15 | 17 | 21 [avatar split clause 574,2343,1687,1647,1599,1547,1543]
2488. 8 | 12 | 15 | 17 | 21 [avatar split clause 573,2343,1687,1647,1599,1573]
2489. 9 | 12 | 15 | 17 | 21 [avatar split clause 572,2343,1687,1647,1599,1578]
2490. ~1 | ~2 | ~3 | 16 | 17 | 21 [avatar split clause 571,2343,1687,1667,1551,1547,1543]
2491. 1 | 2 | ~3 | 16 | 17 | 21 [avatar split clause 570,2343,1687,1667,1551,1547,1543]
2492. 8 | ~3 | 16 | 17 | 21 [avatar split clause 569,2343,1687,1667,1551,1573]
2493. 9 | ~3 | 16 | 17 | 21 [avatar split clause 568,2343,1687,1667,1551,1578]
2494. ~1 | ~2 | 10 | 16 | 17 | 21 [avatar split clause 567,2343,1687,1667,1583,1547,1543]
2495. 1 | 2 | 10 | 16 | 17 | 21 [avatar split clause 566,2343,1687,1667,1583,1547,1543]
2496. 8 | 10 | 16 | 17 | 21 [avatar split clause 565,2343,1687,1667,1583,1573]
2497. 9 | 10 | 16 | 17 | 21 [avatar split clause 564,2343,1687,1667,1583,1578]
2498. ~1 | ~2 | 11 | 16 | 17 | 21 [avatar split clause 563,2343,1687,1667,1591,1547,1543]
2499. 1 | 2 | 11 | 16 | 17 | 21 [avatar split clause 562,2343,1687,1667,1591,1547,1543]
2500. 8 | 11 | 16 | 17 | 21 [avatar split clause 561,2343,1687,1667,1591,1573]
2501. 9 | 11 | 16 | 17 | 21 [avatar split clause 560,2343,1687,1667,1591,1578]
2502. ~1 | ~2 | 12 | 16 | 17 | 21 [avatar split clause 559,2343,1687,1667,1599,1547,1543]
2503. 1 | 2 | 12 | 16 | 17 | 21 [avatar split clause 558,2343,1687,1667,1599,1547,1543]
2504. 8 | 12 | 16 | 17 | 21 [avatar split clause 557,2343,1687,1667,1599,1573]
2505. 9 | 12 | 16 | 17 | 21 [avatar split clause 556,2343,1687,1667,1599,1578]
2506. ~1 | ~2 | ~3 | ~4 | 18 | 21 [avatar split clause 555,2343,1771,1555,1551,1547,1543]
2507. 1 | 2 | ~3 | ~4 | 18 | 21 [avatar split clause 554,2343,1771,1555,1551,1547,1543]
2508. 8 | ~3 | ~4 | 18 | 21 [avatar split clause 553,2343,1771,1555,1551,1573]
2509. 9 | ~3 | ~4 | 18 | 21 [avatar split clause 552,2343,1771,1555,1551,1578]
2510. ~1 | ~2 | 10 | ~4 | 18 | 21 [avatar split clause 551,2343,1771,1555,1583,1547,1543]
2511. 1 | 2 | 10 | ~4 | 18 | 21 [avatar split clause 550,2343,1771,1555,1583,1547,1543]
2512. 8 | 10 | ~4 | 18 | 21 [avatar split clause 549,2343,1771,1555,1583,1573]
2513. 9 | 10 | ~4 | 18 | 21 [avatar split clause 548,2343,1771,1555,1583,1578]
2514. ~1 | ~2 | 11 | ~4 | 18 | 21 [avatar split clause 547,2343,1771,1555,1591,1547,1543]
2515. 1 | 2 | 11 | ~4 | 18 | 21 [avatar split clause 546,2343,1771,1555,1591,1547,1543]
2516. 8 | 11 | ~4 | 18 | 21 [avatar split clause 545,2343,1771,1555,1591,1573]
2517. 9 | 11 | ~4 | 18 | 21 [avatar split clause 544,2343,1771,1555,1591,1578]
2518. ~1 | ~2 | 12 | ~4 | 18 | 21 [avatar split clause 543,2343,1771,1555,1599,1547,1543]
2519. 1 | 2 | 12 | ~4 | 18 | 21 [avatar split clause 542,2343,1771,1555,1599,1547,1543]
2520. 8 | 12 | ~4 | 18 | 21 [avatar split clause 541,2343,1771,1555,1599,1573]
2521. 9 | 12 | ~4 | 18 | 21 [avatar split clause 540,2343,1771,1555,1599,1578]
2522. ~1 | ~2 | ~3 | 13 | 18 | 21 [avatar split clause 539,2343,1771,1607,1551,1547,1543]
2523. 1 | 2 | ~3 | 13 | 18 | 21 [avatar split clause 538,2343,1771,1607,1551,1547,1543]
2524. 8 | ~3 | 13 | 18 | 21 [avatar split clause 537,2343,1771,1607,1551,1573]
2525. 9 | ~3 | 13 | 18 | 21 [avatar split clause 536,2343,1771,1607,1551,1578]
2526. ~1 | ~2 | 10 | 13 | 18 | 21 [avatar split clause 535,2343,1771,1607,1583,1547,1543]
2527. 1 | 2 | 10 | 13 | 18 | 21 [avatar split clause 534,2343,1771,1607,1583,1547,1543]
2528. 8 | 10 | 13 | 18 | 21 [avatar split clause 533,2343,1771,1607,1583,1573]
2529. 9 | 10 | 13 | 18 | 21 [avatar split clause 532,2343,1771,1607,1583,1578]
2530. ~1 | ~2 | 11 | 13 | 18 | 21 [avatar split clause 531,2343,1771,1607,1591,1547,1543]
2531. 1 | 2 | 11 | 13 | 18 | 21 [avatar split clause 530,2343,1771,1607,1591,1547,1543]
2532. 8 | 11 | 13 | 18 | 21 [avatar split clause 529,2343,1771,1607,1591,1573]
2533. 9 | 11 | 13 | 18 | 21 [avatar split clause 528,2343,1771,1607,1591,1578]
2534. ~1 | ~2 | 12 | 13 | 18 | 21 [avatar split clause 527,2343,1771,1607,1599,1547,1543]
2535. 1 | 2 | 12 | 13 | 18 | 21 [avatar split clause 526,2343,1771,1607,1599,1547,1543]
2536. 8 | 12 | 13 | 18 | 21 [avatar split clause 525,2343,1771,1607,1599,1573]
2537. 9 | 12 | 13 | 18 | 21 [avatar split clause 524,2343,1771,1607,1599,1578]
2538. ~1 | ~2 | ~3 | 14 | 18 | 21 [avatar split clause 523,2343,1771,1627,1551,1547,1543]
2539. 1 | 2 | ~3 | 14 | 18 | 21 [avatar split clause 522,2343,1771,1627,1551,1547,1543]
2540. 8 | ~3 | 14 | 18 | 21 [avatar split clause 521,2343,1771,1627,1551,1573]
2541. 9 | ~3 | 14 | 18 | 21 [avatar split clause 520,2343,1771,1627,1551,1578]
2542. ~1 | ~2 | 10 | 14 | 18 | 21 [avatar split clause 519,2343,1771,1627,1583,1547,1543]
2543. 1 | 2 | 10 | 14 | 18 | 21 [avatar split clause 518,2343,1771,1627,1583,1547,1543]
2544. 8 | 10 | 14 | 18 | 21 [avatar split clause 517,2343,1771,1627,1583,1573]
2545. 9 | 10 | 14 | 18 | 21 [avatar split clause 516,2343,1771,1627,1583,1578]
2546. ~1 | ~2 | 11 | 14 | 18 | 21 [avatar split clause 515,2343,1771,1627,1591,1547,1543]
2547. 1 | 2 | 11 | 14 | 18 | 21 [avatar split clause 514,2343,1771,1627,1591,1547,1543]
2548. 8 | 11 | 14 | 18 | 21 [avatar split clause 513,2343,1771,1627,1591,1573]
2549. 9 | 11 | 14 | 18 | 21 [avatar split clause 512,2343,1771,1627,1591,1578]
2550. ~1 | ~2 | 12 | 14 | 18 | 21 [avatar split clause 511,2343,1771,1627,1599,1547,1543]
2551. 1 | 2 | 12 | 14 | 18 | 21 [avatar split clause 510,2343,1771,1627,1599,1547,1543]
2552. 8 | 12 | 14 | 18 | 21 [avatar split clause 509,2343,1771,1627,1599,1573]
2553. 9 | 12 | 14 | 18 | 21 [avatar split clause 508,2343,1771,1627,1599,1578]
2554. ~1 | ~2 | ~3 | 15 | 18 | 21 [avatar split clause 507,2343,1771,1647,1551,1547,1543]
2555. 1 | 2 | ~3 | 15 | 18 | 21 [avatar split clause 506,2343,1771,1647,1551,1547,1543]
2556. 8 | ~3 | 15 | 18 | 21 [avatar split clause 505,2343,1771,1647,1551,1573]
2557. 9 | ~3 | 15 | 18 | 21 [avatar split clause 504,2343,1771,1647,1551,1578]
2558. ~1 | ~2 | 10 | 15 | 18 | 21 [avatar split clause 503,2343,1771,1647,1583,1547,1543]
2559. 1 | 2 | 10 | 15 | 18 | 21 [avatar split clause 502,2343,1771,1647,1583,1547,1543]
2560. 8 | 10 | 15 | 18 | 21 [avatar split clause 501,2343,1771,1647,1583,1573]
2561. 9 | 10 | 15 | 18 | 21 [avatar split clause 500,2343,1771,1647,1583,1578]
2562. ~1 | ~2 | 11 | 15 | 18 | 21 [avatar split clause 499,2343,1771,1647,1591,1547,1543]
2563. 1 | 2 | 11 | 15 | 18 | 21 [avatar split clause 498,2343,1771,1647,1591,1547,1543]
2564. 8 | 11 | 15 | 18 | 21 [avatar split clause 497,2343,1771,1647,1591,1573]
2565. 9 | 11 | 15 | 18 | 21 [avatar split clause 496,2343,1771,1647,1591,1578]
2566. ~1 | ~2 | 12 | 15 | 18 | 21 [avatar split clause 495,2343,1771,1647,1599,1547,1543]
2567. 1 | 2 | 12 | 15 | 18 | 21 [avatar split clause 494,2343,1771,1647,1599,1547,1543]
2568. 8 | 12 | 15 | 18 | 21 [avatar split clause 493,2343,1771,1647,1599,1573]
2569. 9 | 12 | 15 | 18 | 21 [avatar split clause 492,2343,1771,1647,1599,1578]
2570. ~1 | ~2 | ~3 | 16 | 18 | 21 [avatar split clause 491,2343,1771,1667,1551,1547,1543]
2571. 1 | 2 | ~3 | 16 | 18 | 21 [avatar split clause 490,2343,1771,1667,1551,1547,1543]
2572. 8 | ~3 | 16 | 18 | 21 [avatar split clause 489,2343,1771,1667,1551,1573]
2573. 9 | ~3 | 16 | 18 | 21 [avatar split clause 488,2343,1771,1667,1551,1578]
2574. ~1 | ~2 | 10 | 16 | 18 | 21 [avatar split clause 487,2343,1771,1667,1583,1547,1543]
2575. 1 | 2 | 10 | 16 | 18 | 21 [avatar split clause 486,2343,1771,1667,1583,1547,1543]
2576. 8 | 10 | 16 | 18 | 21 [avatar split clause 485,2343,1771,1667,1583,1573]
2577. 9 | 10 | 16 | 18 | 21 [avatar split clause 484,2343,1771,1667,1583,1578]
2578. ~1 | ~2 | 11 | 16 | 18 | 21 [avatar split clause 483,2343,1771,1667,1591,1547,1543]
2579. 1 | 2 | 11 | 16 | 18 | 21 [avatar split clause 482,2343,1771,1667,1591,1547,1543]
2580. 8 | 11 | 16 | 18 | 21 [avatar split clause 481,2343,1771,1667,1591,1573]
2581. 9 | 11 | 16 | 18 | 21 [avatar split clause 480,2343,1771,1667,1591,1578]
2582. ~1 | ~2 | 12 | 16 | 18 | 21 [avatar split clause 479,2343,1771,1667,1599,1547,1543]
2583. 1 | 2 | 12 | 16 | 18 | 21 [avatar split clause 478,2343,1771,1667,1599,1547,1543]
2584. 8 | 12 | 16 | 18 | 21 [avatar split clause 477,2343,1771,1667,1599,1573]
2585. 9 | 12 | 16 | 18 | 21 [avatar split clause 476,2343,1771,1667,1599,1578]
2587. 22 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
2589. m1_subset_1(sK1,u1_struct_0(sK0)) <- (22) [avatar component clause 2587] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2590. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | 22 [avatar split clause 475,2587,1563,1559,1555,1551,1547,1543]
2591. 1 | 2 | ~3 | ~4 | ~5 | ~6 | 22 [avatar split clause 474,2587,1563,1559,1555,1551,1547,1543]
2592. 8 | ~3 | ~4 | ~5 | ~6 | 22 [avatar split clause 473,2587,1563,1559,1555,1551,1573]
2593. 9 | ~3 | ~4 | ~5 | ~6 | 22 [avatar split clause 472,2587,1563,1559,1555,1551,1578]
2594. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | 22 [avatar split clause 471,2587,1563,1559,1555,1583,1547,1543]
2595. 1 | 2 | 10 | ~4 | ~5 | ~6 | 22 [avatar split clause 470,2587,1563,1559,1555,1583,1547,1543]
2596. 8 | 10 | ~4 | ~5 | ~6 | 22 [avatar split clause 469,2587,1563,1559,1555,1583,1573]
2597. 9 | 10 | ~4 | ~5 | ~6 | 22 [avatar split clause 468,2587,1563,1559,1555,1583,1578]
2598. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | 22 [avatar split clause 467,2587,1563,1559,1555,1591,1547,1543]
2599. 1 | 2 | 11 | ~4 | ~5 | ~6 | 22 [avatar split clause 466,2587,1563,1559,1555,1591,1547,1543]
2600. 8 | 11 | ~4 | ~5 | ~6 | 22 [avatar split clause 465,2587,1563,1559,1555,1591,1573]
2601. 9 | 11 | ~4 | ~5 | ~6 | 22 [avatar split clause 464,2587,1563,1559,1555,1591,1578]
2602. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | 22 [avatar split clause 463,2587,1563,1559,1555,1599,1547,1543]
2603. 1 | 2 | 12 | ~4 | ~5 | ~6 | 22 [avatar split clause 462,2587,1563,1559,1555,1599,1547,1543]
2604. 8 | 12 | ~4 | ~5 | ~6 | 22 [avatar split clause 461,2587,1563,1559,1555,1599,1573]
2605. 9 | 12 | ~4 | ~5 | ~6 | 22 [avatar split clause 460,2587,1563,1559,1555,1599,1578]
2606. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | 22 [avatar split clause 459,2587,1563,1559,1607,1551,1547,1543]
2607. 1 | 2 | ~3 | 13 | ~5 | ~6 | 22 [avatar split clause 458,2587,1563,1559,1607,1551,1547,1543]
2608. 8 | ~3 | 13 | ~5 | ~6 | 22 [avatar split clause 457,2587,1563,1559,1607,1551,1573]
2609. 9 | ~3 | 13 | ~5 | ~6 | 22 [avatar split clause 456,2587,1563,1559,1607,1551,1578]
2610. ~1 | ~2 | 10 | 13 | ~5 | ~6 | 22 [avatar split clause 455,2587,1563,1559,1607,1583,1547,1543]
2611. 1 | 2 | 10 | 13 | ~5 | ~6 | 22 [avatar split clause 454,2587,1563,1559,1607,1583,1547,1543]
2612. 8 | 10 | 13 | ~5 | ~6 | 22 [avatar split clause 453,2587,1563,1559,1607,1583,1573]
2613. 9 | 10 | 13 | ~5 | ~6 | 22 [avatar split clause 452,2587,1563,1559,1607,1583,1578]
2614. ~1 | ~2 | 11 | 13 | ~5 | ~6 | 22 [avatar split clause 451,2587,1563,1559,1607,1591,1547,1543]
2615. 1 | 2 | 11 | 13 | ~5 | ~6 | 22 [avatar split clause 450,2587,1563,1559,1607,1591,1547,1543]
2616. 8 | 11 | 13 | ~5 | ~6 | 22 [avatar split clause 449,2587,1563,1559,1607,1591,1573]
2617. 9 | 11 | 13 | ~5 | ~6 | 22 [avatar split clause 448,2587,1563,1559,1607,1591,1578]
2618. ~1 | ~2 | 12 | 13 | ~5 | ~6 | 22 [avatar split clause 447,2587,1563,1559,1607,1599,1547,1543]
2619. 1 | 2 | 12 | 13 | ~5 | ~6 | 22 [avatar split clause 446,2587,1563,1559,1607,1599,1547,1543]
2620. 8 | 12 | 13 | ~5 | ~6 | 22 [avatar split clause 445,2587,1563,1559,1607,1599,1573]
2621. 9 | 12 | 13 | ~5 | ~6 | 22 [avatar split clause 444,2587,1563,1559,1607,1599,1578]
2622. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | 22 [avatar split clause 443,2587,1563,1559,1627,1551,1547,1543]
2623. 1 | 2 | ~3 | 14 | ~5 | ~6 | 22 [avatar split clause 442,2587,1563,1559,1627,1551,1547,1543]
2624. 8 | ~3 | 14 | ~5 | ~6 | 22 [avatar split clause 441,2587,1563,1559,1627,1551,1573]
2625. 9 | ~3 | 14 | ~5 | ~6 | 22 [avatar split clause 440,2587,1563,1559,1627,1551,1578]
2626. ~1 | ~2 | 10 | 14 | ~5 | ~6 | 22 [avatar split clause 439,2587,1563,1559,1627,1583,1547,1543]
2627. 1 | 2 | 10 | 14 | ~5 | ~6 | 22 [avatar split clause 438,2587,1563,1559,1627,1583,1547,1543]
2628. 8 | 10 | 14 | ~5 | ~6 | 22 [avatar split clause 437,2587,1563,1559,1627,1583,1573]
2629. 9 | 10 | 14 | ~5 | ~6 | 22 [avatar split clause 436,2587,1563,1559,1627,1583,1578]
2630. ~1 | ~2 | 11 | 14 | ~5 | ~6 | 22 [avatar split clause 435,2587,1563,1559,1627,1591,1547,1543]
2631. 1 | 2 | 11 | 14 | ~5 | ~6 | 22 [avatar split clause 434,2587,1563,1559,1627,1591,1547,1543]
2632. 8 | 11 | 14 | ~5 | ~6 | 22 [avatar split clause 433,2587,1563,1559,1627,1591,1573]
2633. 9 | 11 | 14 | ~5 | ~6 | 22 [avatar split clause 432,2587,1563,1559,1627,1591,1578]
2634. ~1 | ~2 | 12 | 14 | ~5 | ~6 | 22 [avatar split clause 431,2587,1563,1559,1627,1599,1547,1543]
2635. 1 | 2 | 12 | 14 | ~5 | ~6 | 22 [avatar split clause 430,2587,1563,1559,1627,1599,1547,1543]
2636. 8 | 12 | 14 | ~5 | ~6 | 22 [avatar split clause 429,2587,1563,1559,1627,1599,1573]
2637. 9 | 12 | 14 | ~5 | ~6 | 22 [avatar split clause 428,2587,1563,1559,1627,1599,1578]
2638. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | 22 [avatar split clause 427,2587,1563,1559,1647,1551,1547,1543]
2639. 1 | 2 | ~3 | 15 | ~5 | ~6 | 22 [avatar split clause 426,2587,1563,1559,1647,1551,1547,1543]
2640. 8 | ~3 | 15 | ~5 | ~6 | 22 [avatar split clause 425,2587,1563,1559,1647,1551,1573]
2641. 9 | ~3 | 15 | ~5 | ~6 | 22 [avatar split clause 424,2587,1563,1559,1647,1551,1578]
2642. ~1 | ~2 | 10 | 15 | ~5 | ~6 | 22 [avatar split clause 423,2587,1563,1559,1647,1583,1547,1543]
2643. 1 | 2 | 10 | 15 | ~5 | ~6 | 22 [avatar split clause 422,2587,1563,1559,1647,1583,1547,1543]
2644. 8 | 10 | 15 | ~5 | ~6 | 22 [avatar split clause 421,2587,1563,1559,1647,1583,1573]
2645. 9 | 10 | 15 | ~5 | ~6 | 22 [avatar split clause 420,2587,1563,1559,1647,1583,1578]
2646. ~1 | ~2 | 11 | 15 | ~5 | ~6 | 22 [avatar split clause 419,2587,1563,1559,1647,1591,1547,1543]
2647. 1 | 2 | 11 | 15 | ~5 | ~6 | 22 [avatar split clause 418,2587,1563,1559,1647,1591,1547,1543]
2648. 8 | 11 | 15 | ~5 | ~6 | 22 [avatar split clause 417,2587,1563,1559,1647,1591,1573]
2649. 9 | 11 | 15 | ~5 | ~6 | 22 [avatar split clause 416,2587,1563,1559,1647,1591,1578]
2650. ~1 | ~2 | 12 | 15 | ~5 | ~6 | 22 [avatar split clause 415,2587,1563,1559,1647,1599,1547,1543]
2651. 1 | 2 | 12 | 15 | ~5 | ~6 | 22 [avatar split clause 414,2587,1563,1559,1647,1599,1547,1543]
2652. 8 | 12 | 15 | ~5 | ~6 | 22 [avatar split clause 413,2587,1563,1559,1647,1599,1573]
2653. 9 | 12 | 15 | ~5 | ~6 | 22 [avatar split clause 412,2587,1563,1559,1647,1599,1578]
2654. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | 22 [avatar split clause 411,2587,1563,1559,1667,1551,1547,1543]
2655. 1 | 2 | ~3 | 16 | ~5 | ~6 | 22 [avatar split clause 410,2587,1563,1559,1667,1551,1547,1543]
2656. 8 | ~3 | 16 | ~5 | ~6 | 22 [avatar split clause 409,2587,1563,1559,1667,1551,1573]
2657. 9 | ~3 | 16 | ~5 | ~6 | 22 [avatar split clause 408,2587,1563,1559,1667,1551,1578]
2658. ~1 | ~2 | 10 | 16 | ~5 | ~6 | 22 [avatar split clause 407,2587,1563,1559,1667,1583,1547,1543]
2659. 1 | 2 | 10 | 16 | ~5 | ~6 | 22 [avatar split clause 406,2587,1563,1559,1667,1583,1547,1543]
2660. 8 | 10 | 16 | ~5 | ~6 | 22 [avatar split clause 405,2587,1563,1559,1667,1583,1573]
2661. 9 | 10 | 16 | ~5 | ~6 | 22 [avatar split clause 404,2587,1563,1559,1667,1583,1578]
2662. ~1 | ~2 | 11 | 16 | ~5 | ~6 | 22 [avatar split clause 403,2587,1563,1559,1667,1591,1547,1543]
2663. 1 | 2 | 11 | 16 | ~5 | ~6 | 22 [avatar split clause 402,2587,1563,1559,1667,1591,1547,1543]
2664. 8 | 11 | 16 | ~5 | ~6 | 22 [avatar split clause 401,2587,1563,1559,1667,1591,1573]
2665. 9 | 11 | 16 | ~5 | ~6 | 22 [avatar split clause 400,2587,1563,1559,1667,1591,1578]
2666. ~1 | ~2 | 12 | 16 | ~5 | ~6 | 22 [avatar split clause 399,2587,1563,1559,1667,1599,1547,1543]
2667. 1 | 2 | 12 | 16 | ~5 | ~6 | 22 [avatar split clause 398,2587,1563,1559,1667,1599,1547,1543]
2668. 8 | 12 | 16 | ~5 | ~6 | 22 [avatar split clause 397,2587,1563,1559,1667,1599,1573]
2669. 9 | 12 | 16 | ~5 | ~6 | 22 [avatar split clause 396,2587,1563,1559,1667,1599,1578]
2670. ~1 | ~2 | ~3 | ~4 | 17 | 22 [avatar split clause 395,2587,1687,1555,1551,1547,1543]
2671. 1 | 2 | ~3 | ~4 | 17 | 22 [avatar split clause 394,2587,1687,1555,1551,1547,1543]
2672. 8 | ~3 | ~4 | 17 | 22 [avatar split clause 393,2587,1687,1555,1551,1573]
2673. 9 | ~3 | ~4 | 17 | 22 [avatar split clause 392,2587,1687,1555,1551,1578]
2674. ~1 | ~2 | 10 | ~4 | 17 | 22 [avatar split clause 391,2587,1687,1555,1583,1547,1543]
2675. 1 | 2 | 10 | ~4 | 17 | 22 [avatar split clause 390,2587,1687,1555,1583,1547,1543]
2676. 8 | 10 | ~4 | 17 | 22 [avatar split clause 389,2587,1687,1555,1583,1573]
2677. 9 | 10 | ~4 | 17 | 22 [avatar split clause 388,2587,1687,1555,1583,1578]
2678. ~1 | ~2 | 11 | ~4 | 17 | 22 [avatar split clause 387,2587,1687,1555,1591,1547,1543]
2679. 1 | 2 | 11 | ~4 | 17 | 22 [avatar split clause 386,2587,1687,1555,1591,1547,1543]
2680. 8 | 11 | ~4 | 17 | 22 [avatar split clause 385,2587,1687,1555,1591,1573]
2681. 9 | 11 | ~4 | 17 | 22 [avatar split clause 384,2587,1687,1555,1591,1578]
2682. ~1 | ~2 | 12 | ~4 | 17 | 22 [avatar split clause 383,2587,1687,1555,1599,1547,1543]
2683. 1 | 2 | 12 | ~4 | 17 | 22 [avatar split clause 382,2587,1687,1555,1599,1547,1543]
2684. 8 | 12 | ~4 | 17 | 22 [avatar split clause 381,2587,1687,1555,1599,1573]
2685. 9 | 12 | ~4 | 17 | 22 [avatar split clause 380,2587,1687,1555,1599,1578]
2686. ~1 | ~2 | ~3 | 13 | 17 | 22 [avatar split clause 379,2587,1687,1607,1551,1547,1543]
2687. 1 | 2 | ~3 | 13 | 17 | 22 [avatar split clause 378,2587,1687,1607,1551,1547,1543]
2688. 8 | ~3 | 13 | 17 | 22 [avatar split clause 377,2587,1687,1607,1551,1573]
2689. 9 | ~3 | 13 | 17 | 22 [avatar split clause 376,2587,1687,1607,1551,1578]
2690. ~1 | ~2 | 10 | 13 | 17 | 22 [avatar split clause 375,2587,1687,1607,1583,1547,1543]
2691. 1 | 2 | 10 | 13 | 17 | 22 [avatar split clause 374,2587,1687,1607,1583,1547,1543]
2692. 8 | 10 | 13 | 17 | 22 [avatar split clause 373,2587,1687,1607,1583,1573]
2693. 9 | 10 | 13 | 17 | 22 [avatar split clause 372,2587,1687,1607,1583,1578]
2694. ~1 | ~2 | 11 | 13 | 17 | 22 [avatar split clause 371,2587,1687,1607,1591,1547,1543]
2695. 1 | 2 | 11 | 13 | 17 | 22 [avatar split clause 370,2587,1687,1607,1591,1547,1543]
2696. 8 | 11 | 13 | 17 | 22 [avatar split clause 369,2587,1687,1607,1591,1573]
2697. 9 | 11 | 13 | 17 | 22 [avatar split clause 368,2587,1687,1607,1591,1578]
2698. ~1 | ~2 | 12 | 13 | 17 | 22 [avatar split clause 367,2587,1687,1607,1599,1547,1543]
2699. 1 | 2 | 12 | 13 | 17 | 22 [avatar split clause 366,2587,1687,1607,1599,1547,1543]
2700. 8 | 12 | 13 | 17 | 22 [avatar split clause 365,2587,1687,1607,1599,1573]
2701. 9 | 12 | 13 | 17 | 22 [avatar split clause 364,2587,1687,1607,1599,1578]
2702. ~1 | ~2 | ~3 | 14 | 17 | 22 [avatar split clause 363,2587,1687,1627,1551,1547,1543]
2703. 1 | 2 | ~3 | 14 | 17 | 22 [avatar split clause 362,2587,1687,1627,1551,1547,1543]
2704. 8 | ~3 | 14 | 17 | 22 [avatar split clause 361,2587,1687,1627,1551,1573]
2705. 9 | ~3 | 14 | 17 | 22 [avatar split clause 360,2587,1687,1627,1551,1578]
2706. ~1 | ~2 | 10 | 14 | 17 | 22 [avatar split clause 359,2587,1687,1627,1583,1547,1543]
2707. 1 | 2 | 10 | 14 | 17 | 22 [avatar split clause 358,2587,1687,1627,1583,1547,1543]
2708. 8 | 10 | 14 | 17 | 22 [avatar split clause 357,2587,1687,1627,1583,1573]
2709. 9 | 10 | 14 | 17 | 22 [avatar split clause 356,2587,1687,1627,1583,1578]
2710. ~1 | ~2 | 11 | 14 | 17 | 22 [avatar split clause 355,2587,1687,1627,1591,1547,1543]
2711. 1 | 2 | 11 | 14 | 17 | 22 [avatar split clause 354,2587,1687,1627,1591,1547,1543]
2712. 8 | 11 | 14 | 17 | 22 [avatar split clause 353,2587,1687,1627,1591,1573]
2713. 9 | 11 | 14 | 17 | 22 [avatar split clause 352,2587,1687,1627,1591,1578]
2714. ~1 | ~2 | 12 | 14 | 17 | 22 [avatar split clause 351,2587,1687,1627,1599,1547,1543]
2715. 1 | 2 | 12 | 14 | 17 | 22 [avatar split clause 350,2587,1687,1627,1599,1547,1543]
2716. 8 | 12 | 14 | 17 | 22 [avatar split clause 349,2587,1687,1627,1599,1573]
2717. 9 | 12 | 14 | 17 | 22 [avatar split clause 348,2587,1687,1627,1599,1578]
2718. ~1 | ~2 | ~3 | 15 | 17 | 22 [avatar split clause 347,2587,1687,1647,1551,1547,1543]
2719. 1 | 2 | ~3 | 15 | 17 | 22 [avatar split clause 346,2587,1687,1647,1551,1547,1543]
2720. 8 | ~3 | 15 | 17 | 22 [avatar split clause 345,2587,1687,1647,1551,1573]
2721. 9 | ~3 | 15 | 17 | 22 [avatar split clause 344,2587,1687,1647,1551,1578]
2722. ~1 | ~2 | 10 | 15 | 17 | 22 [avatar split clause 343,2587,1687,1647,1583,1547,1543]
2723. 1 | 2 | 10 | 15 | 17 | 22 [avatar split clause 342,2587,1687,1647,1583,1547,1543]
2724. 8 | 10 | 15 | 17 | 22 [avatar split clause 341,2587,1687,1647,1583,1573]
2725. 9 | 10 | 15 | 17 | 22 [avatar split clause 340,2587,1687,1647,1583,1578]
2726. ~1 | ~2 | 11 | 15 | 17 | 22 [avatar split clause 339,2587,1687,1647,1591,1547,1543]
2727. 1 | 2 | 11 | 15 | 17 | 22 [avatar split clause 338,2587,1687,1647,1591,1547,1543]
2728. 8 | 11 | 15 | 17 | 22 [avatar split clause 337,2587,1687,1647,1591,1573]
2729. 9 | 11 | 15 | 17 | 22 [avatar split clause 336,2587,1687,1647,1591,1578]
2730. ~1 | ~2 | 12 | 15 | 17 | 22 [avatar split clause 335,2587,1687,1647,1599,1547,1543]
2731. 1 | 2 | 12 | 15 | 17 | 22 [avatar split clause 334,2587,1687,1647,1599,1547,1543]
2732. 8 | 12 | 15 | 17 | 22 [avatar split clause 333,2587,1687,1647,1599,1573]
2733. 9 | 12 | 15 | 17 | 22 [avatar split clause 332,2587,1687,1647,1599,1578]
2734. ~1 | ~2 | ~3 | 16 | 17 | 22 [avatar split clause 331,2587,1687,1667,1551,1547,1543]
2735. 1 | 2 | ~3 | 16 | 17 | 22 [avatar split clause 330,2587,1687,1667,1551,1547,1543]
2736. 8 | ~3 | 16 | 17 | 22 [avatar split clause 329,2587,1687,1667,1551,1573]
2737. 9 | ~3 | 16 | 17 | 22 [avatar split clause 328,2587,1687,1667,1551,1578]
2738. ~1 | ~2 | 10 | 16 | 17 | 22 [avatar split clause 327,2587,1687,1667,1583,1547,1543]
2739. 1 | 2 | 10 | 16 | 17 | 22 [avatar split clause 326,2587,1687,1667,1583,1547,1543]
2740. 8 | 10 | 16 | 17 | 22 [avatar split clause 325,2587,1687,1667,1583,1573]
2741. 9 | 10 | 16 | 17 | 22 [avatar split clause 324,2587,1687,1667,1583,1578]
2742. ~1 | ~2 | 11 | 16 | 17 | 22 [avatar split clause 323,2587,1687,1667,1591,1547,1543]
2743. 1 | 2 | 11 | 16 | 17 | 22 [avatar split clause 322,2587,1687,1667,1591,1547,1543]
2744. 8 | 11 | 16 | 17 | 22 [avatar split clause 321,2587,1687,1667,1591,1573]
2745. 9 | 11 | 16 | 17 | 22 [avatar split clause 320,2587,1687,1667,1591,1578]
2746. ~1 | ~2 | 12 | 16 | 17 | 22 [avatar split clause 319,2587,1687,1667,1599,1547,1543]
2747. 1 | 2 | 12 | 16 | 17 | 22 [avatar split clause 318,2587,1687,1667,1599,1547,1543]
2748. 8 | 12 | 16 | 17 | 22 [avatar split clause 317,2587,1687,1667,1599,1573]
2749. 9 | 12 | 16 | 17 | 22 [avatar split clause 316,2587,1687,1667,1599,1578]
2750. ~1 | ~2 | ~3 | ~4 | 18 | 22 [avatar split clause 315,2587,1771,1555,1551,1547,1543]
2751. 1 | 2 | ~3 | ~4 | 18 | 22 [avatar split clause 314,2587,1771,1555,1551,1547,1543]
2752. 8 | ~3 | ~4 | 18 | 22 [avatar split clause 313,2587,1771,1555,1551,1573]
2753. 9 | ~3 | ~4 | 18 | 22 [avatar split clause 312,2587,1771,1555,1551,1578]
2754. ~1 | ~2 | 10 | ~4 | 18 | 22 [avatar split clause 311,2587,1771,1555,1583,1547,1543]
2755. 1 | 2 | 10 | ~4 | 18 | 22 [avatar split clause 310,2587,1771,1555,1583,1547,1543]
2756. 8 | 10 | ~4 | 18 | 22 [avatar split clause 309,2587,1771,1555,1583,1573]
2757. 9 | 10 | ~4 | 18 | 22 [avatar split clause 308,2587,1771,1555,1583,1578]
2758. ~1 | ~2 | 11 | ~4 | 18 | 22 [avatar split clause 307,2587,1771,1555,1591,1547,1543]
2759. 1 | 2 | 11 | ~4 | 18 | 22 [avatar split clause 306,2587,1771,1555,1591,1547,1543]
2760. 8 | 11 | ~4 | 18 | 22 [avatar split clause 305,2587,1771,1555,1591,1573]
2761. 9 | 11 | ~4 | 18 | 22 [avatar split clause 304,2587,1771,1555,1591,1578]
2762. ~1 | ~2 | 12 | ~4 | 18 | 22 [avatar split clause 303,2587,1771,1555,1599,1547,1543]
2763. 1 | 2 | 12 | ~4 | 18 | 22 [avatar split clause 302,2587,1771,1555,1599,1547,1543]
2764. 8 | 12 | ~4 | 18 | 22 [avatar split clause 301,2587,1771,1555,1599,1573]
2765. 9 | 12 | ~4 | 18 | 22 [avatar split clause 300,2587,1771,1555,1599,1578]
2766. ~1 | ~2 | ~3 | 13 | 18 | 22 [avatar split clause 299,2587,1771,1607,1551,1547,1543]
2767. 1 | 2 | ~3 | 13 | 18 | 22 [avatar split clause 298,2587,1771,1607,1551,1547,1543]
2768. 8 | ~3 | 13 | 18 | 22 [avatar split clause 297,2587,1771,1607,1551,1573]
2769. 9 | ~3 | 13 | 18 | 22 [avatar split clause 296,2587,1771,1607,1551,1578]
2770. ~1 | ~2 | 10 | 13 | 18 | 22 [avatar split clause 295,2587,1771,1607,1583,1547,1543]
2771. 1 | 2 | 10 | 13 | 18 | 22 [avatar split clause 294,2587,1771,1607,1583,1547,1543]
2772. 8 | 10 | 13 | 18 | 22 [avatar split clause 293,2587,1771,1607,1583,1573]
2773. 9 | 10 | 13 | 18 | 22 [avatar split clause 292,2587,1771,1607,1583,1578]
2774. ~1 | ~2 | 11 | 13 | 18 | 22 [avatar split clause 291,2587,1771,1607,1591,1547,1543]
2775. 1 | 2 | 11 | 13 | 18 | 22 [avatar split clause 290,2587,1771,1607,1591,1547,1543]
2776. 8 | 11 | 13 | 18 | 22 [avatar split clause 289,2587,1771,1607,1591,1573]
2777. 9 | 11 | 13 | 18 | 22 [avatar split clause 288,2587,1771,1607,1591,1578]
2778. ~1 | ~2 | 12 | 13 | 18 | 22 [avatar split clause 287,2587,1771,1607,1599,1547,1543]
2779. 1 | 2 | 12 | 13 | 18 | 22 [avatar split clause 286,2587,1771,1607,1599,1547,1543]
2780. 8 | 12 | 13 | 18 | 22 [avatar split clause 285,2587,1771,1607,1599,1573]
2781. 9 | 12 | 13 | 18 | 22 [avatar split clause 284,2587,1771,1607,1599,1578]
2782. ~1 | ~2 | ~3 | 14 | 18 | 22 [avatar split clause 283,2587,1771,1627,1551,1547,1543]
2783. 1 | 2 | ~3 | 14 | 18 | 22 [avatar split clause 282,2587,1771,1627,1551,1547,1543]
2784. 8 | ~3 | 14 | 18 | 22 [avatar split clause 281,2587,1771,1627,1551,1573]
2785. 9 | ~3 | 14 | 18 | 22 [avatar split clause 280,2587,1771,1627,1551,1578]
2786. ~1 | ~2 | 10 | 14 | 18 | 22 [avatar split clause 279,2587,1771,1627,1583,1547,1543]
2787. 1 | 2 | 10 | 14 | 18 | 22 [avatar split clause 278,2587,1771,1627,1583,1547,1543]
2788. 8 | 10 | 14 | 18 | 22 [avatar split clause 277,2587,1771,1627,1583,1573]
2789. 9 | 10 | 14 | 18 | 22 [avatar split clause 276,2587,1771,1627,1583,1578]
2790. ~1 | ~2 | 11 | 14 | 18 | 22 [avatar split clause 275,2587,1771,1627,1591,1547,1543]
2791. 1 | 2 | 11 | 14 | 18 | 22 [avatar split clause 274,2587,1771,1627,1591,1547,1543]
2792. 8 | 11 | 14 | 18 | 22 [avatar split clause 273,2587,1771,1627,1591,1573]
2793. 9 | 11 | 14 | 18 | 22 [avatar split clause 272,2587,1771,1627,1591,1578]
2794. ~1 | ~2 | 12 | 14 | 18 | 22 [avatar split clause 271,2587,1771,1627,1599,1547,1543]
2795. 1 | 2 | 12 | 14 | 18 | 22 [avatar split clause 270,2587,1771,1627,1599,1547,1543]
2796. 8 | 12 | 14 | 18 | 22 [avatar split clause 269,2587,1771,1627,1599,1573]
2797. 9 | 12 | 14 | 18 | 22 [avatar split clause 268,2587,1771,1627,1599,1578]
2798. ~1 | ~2 | ~3 | 15 | 18 | 22 [avatar split clause 267,2587,1771,1647,1551,1547,1543]
2799. 1 | 2 | ~3 | 15 | 18 | 22 [avatar split clause 266,2587,1771,1647,1551,1547,1543]
2800. 8 | ~3 | 15 | 18 | 22 [avatar split clause 265,2587,1771,1647,1551,1573]
2801. 9 | ~3 | 15 | 18 | 22 [avatar split clause 264,2587,1771,1647,1551,1578]
2802. ~1 | ~2 | 10 | 15 | 18 | 22 [avatar split clause 263,2587,1771,1647,1583,1547,1543]
2803. 1 | 2 | 10 | 15 | 18 | 22 [avatar split clause 262,2587,1771,1647,1583,1547,1543]
2804. 8 | 10 | 15 | 18 | 22 [avatar split clause 261,2587,1771,1647,1583,1573]
2805. 9 | 10 | 15 | 18 | 22 [avatar split clause 260,2587,1771,1647,1583,1578]
2806. ~1 | ~2 | 11 | 15 | 18 | 22 [avatar split clause 259,2587,1771,1647,1591,1547,1543]
2807. 1 | 2 | 11 | 15 | 18 | 22 [avatar split clause 258,2587,1771,1647,1591,1547,1543]
2808. 8 | 11 | 15 | 18 | 22 [avatar split clause 257,2587,1771,1647,1591,1573]
2809. 9 | 11 | 15 | 18 | 22 [avatar split clause 256,2587,1771,1647,1591,1578]
2810. ~1 | ~2 | 12 | 15 | 18 | 22 [avatar split clause 255,2587,1771,1647,1599,1547,1543]
2811. 1 | 2 | 12 | 15 | 18 | 22 [avatar split clause 254,2587,1771,1647,1599,1547,1543]
2812. 8 | 12 | 15 | 18 | 22 [avatar split clause 253,2587,1771,1647,1599,1573]
2813. 9 | 12 | 15 | 18 | 22 [avatar split clause 252,2587,1771,1647,1599,1578]
2814. ~1 | ~2 | ~3 | 16 | 18 | 22 [avatar split clause 251,2587,1771,1667,1551,1547,1543]
2815. 1 | 2 | ~3 | 16 | 18 | 22 [avatar split clause 250,2587,1771,1667,1551,1547,1543]
2816. 8 | ~3 | 16 | 18 | 22 [avatar split clause 249,2587,1771,1667,1551,1573]
2817. 9 | ~3 | 16 | 18 | 22 [avatar split clause 248,2587,1771,1667,1551,1578]
2818. ~1 | ~2 | 10 | 16 | 18 | 22 [avatar split clause 247,2587,1771,1667,1583,1547,1543]
2819. 1 | 2 | 10 | 16 | 18 | 22 [avatar split clause 246,2587,1771,1667,1583,1547,1543]
2820. 8 | 10 | 16 | 18 | 22 [avatar split clause 245,2587,1771,1667,1583,1573]
2821. 9 | 10 | 16 | 18 | 22 [avatar split clause 244,2587,1771,1667,1583,1578]
2822. ~1 | ~2 | 11 | 16 | 18 | 22 [avatar split clause 243,2587,1771,1667,1591,1547,1543]
2823. 1 | 2 | 11 | 16 | 18 | 22 [avatar split clause 242,2587,1771,1667,1591,1547,1543]
2824. 8 | 11 | 16 | 18 | 22 [avatar split clause 241,2587,1771,1667,1591,1573]
2825. 9 | 11 | 16 | 18 | 22 [avatar split clause 240,2587,1771,1667,1591,1578]
2826. ~1 | ~2 | 12 | 16 | 18 | 22 [avatar split clause 239,2587,1771,1667,1599,1547,1543]
2827. 1 | 2 | 12 | 16 | 18 | 22 [avatar split clause 238,2587,1771,1667,1599,1547,1543]
2828. 8 | 12 | 16 | 18 | 22 [avatar split clause 237,2587,1771,1667,1599,1573]
2829. 9 | 12 | 16 | 18 | 22 [avatar split clause 236,2587,1771,1667,1599,1578]
2831. 23 <=> m2_bcialg_1(k1_tarski(k4_struct_0(sK0)),sK0) [avatar definition]
2833. m2_bcialg_1(k1_tarski(k4_struct_0(sK0)),sK0) <- (23) [avatar component clause 2831] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2834. 23 [avatar split clause 235,2831]
2836. 24 <=> v3_bciideal(k1_tarski(k4_struct_0(sK0)),sK0) [avatar definition]
2838. v3_bciideal(k1_tarski(k4_struct_0(sK0)),sK0) <- (24) [avatar component clause 2836] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2839. 24 [avatar split clause 234,2836]
2841. 25 <=> l2_bcialg_1(sK0) [avatar definition]
2843. l2_bcialg_1(sK0) <- (25) [avatar component clause 2841] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2844. 25 [avatar split clause 233,2841]
2846. 26 <=> v8_bcialg_1(sK0) [avatar definition]
2848. v8_bcialg_1(sK0) <- (26) [avatar component clause 2846] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2849. 26 [avatar split clause 232,2846]
2851. 27 <=> v7_bcialg_1(sK0) [avatar definition]
2853. v7_bcialg_1(sK0) <- (27) [avatar component clause 2851] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2854. 27 [avatar split clause 231,2851]
2856. 28 <=> v5_bcialg_1(sK0) [avatar definition]
2858. v5_bcialg_1(sK0) <- (28) [avatar component clause 2856] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2859. 28 [avatar split clause 230,2856]
2861. 29 <=> v4_bcialg_1(sK0) [avatar definition]
2863. v4_bcialg_1(sK0) <- (29) [avatar component clause 2861] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2864. 29 [avatar split clause 229,2861]
2866. 30 <=> v3_bcialg_1(sK0) [avatar definition]
2868. v3_bcialg_1(sK0) <- (30) [avatar component clause 2866] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2869. 30 [avatar split clause 228,2866]
2871. 31 <=> v2_struct_0(sK0) [avatar definition]
2873. ~v2_struct_0(sK0) <- (~31) [avatar component clause 2871] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2874. ~31 [avatar split clause 227,2871]
3304. v1_bcialg_3(sK0) | ~v3_bciideal(k1_tarski(k4_struct_0(sK0)),sK0) | ~l2_bcialg_1(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23) [resolution 1485,2833] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3307. v1_bcialg_3(sK0) | ~l2_bcialg_1(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24) [subsumption resolution 3304,2838] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3308. v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24, 25) [subsumption resolution 3307,2843] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3309. v1_bcialg_3(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24, 25, 26) [subsumption resolution 3308,2848] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3310. v1_bcialg_3(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24, 25, 26, 27) [subsumption resolution 3309,2853] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3311. v1_bcialg_3(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24, 25, 26, 27, 28) [subsumption resolution 3310,2858] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3312. v1_bcialg_3(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24, 25, 26, 27, 28, 29) [subsumption resolution 3311,2863] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3313. v1_bcialg_3(sK0) | v2_struct_0(sK0) <- (23, 24, 25, 26, 27, 28, 29, 30) [subsumption resolution 3312,2868] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3314. v1_bcialg_3(sK0) <- (23, 24, 25, 26, 27, 28, 29, 30, ~31) [subsumption resolution 3313,2873] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3316. 90 <=> v1_bcialg_3(sK0) [avatar definition]
3318. v1_bcialg_3(sK0) <- (90) [avatar component clause 3316] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:15,thDist:-15}
3319. 90 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | ~29 | ~30 | 31 [avatar split clause 3314,2871,2866,2861,2856,2851,2846,2841,2836,2831,3316]
3336. sK9 != sK9 | sK10 = k1_bcialg_1(sK0,sK10,sK9) | ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (10) [superposition 1539,1585] {a:1,w:33,wCS:3300,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3338. sK10 = k1_bcialg_1(sK0,sK10,sK9) | ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (10) [trivial inequality removal 3336] {a:1,w:30,wCS:3000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3347. k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,sK5) | ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5) [superposition 1561,1540] {a:1,w:33,wCS:3300,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3352. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5) [trivial inequality removal 3347] {a:1,w:24,wCS:2400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3366. k4_struct_0(sK0) != k4_struct_0(sK0) | sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2) [superposition 1538,1548] {a:1,w:35,wCS:3500,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3369. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2) [trivial inequality removal 3366] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3410. k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,sK7,sK6) | ~r1_bcialg_1(sK0,sK7,sK8) | ~m1_subset_1(sK8,u1_struct_0(sK0)) | ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~m1_subset_1(sK6,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4) [superposition 1557,1536] {a:1,w:41,wCS:4100,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3424. ~r1_bcialg_1(sK0,sK7,sK8) | ~m1_subset_1(sK8,u1_struct_0(sK0)) | ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~m1_subset_1(sK6,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4) [trivial inequality removal 3410] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3433. ~m1_subset_1(sK8,u1_struct_0(sK0)) | ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~m1_subset_1(sK6,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13) [subsumption resolution 3424,1609] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3434. ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~m1_subset_1(sK6,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14) [subsumption resolution 3433,1629] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3446. ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 16) [subsumption resolution 3434,1669] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3447. ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 16, 25) [subsumption resolution 3446,2843] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3448. ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 16, 25, 90) [subsumption resolution 3447,3318] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3449. ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 16, 25, 26, 90) [subsumption resolution 3448,2848] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3450. ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 16, 25, 26, 27, 90) [subsumption resolution 3449,2853] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3451. ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 16, 25, 26, 27, 28, 90) [subsumption resolution 3450,2858] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3452. ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 16, 25, 26, 27, 28, 29, 90) [subsumption resolution 3451,2863] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3453. ~m1_subset_1(sK7,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~4, 13, 14, 16, 25, 26, 27, 28, 29, 30, 90) [subsumption resolution 3452,2868] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3454. ~m1_subset_1(sK7,u1_struct_0(sK0)) <- (~4, 13, 14, 16, 25, 26, 27, 28, 29, 30, ~31, 90) [subsumption resolution 3453,2873] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3455. ~15 | 4 | ~13 | ~14 | ~16 | ~25 | ~26 | ~27 | ~28 | ~29 | ~30 | 31 | ~90 [avatar split clause 3454,3316,2871,2866,2861,2856,2851,2846,2841,1667,1627,1607,1555,1647]
3457. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 25) [subsumption resolution 3352,2843] {a:1,w:22,wCS:2200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3458. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 25, 90) [subsumption resolution 3457,3318] {a:1,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3459. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 25, 26, 90) [subsumption resolution 3458,2848] {a:1,w:18,wCS:1800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3460. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 25, 26, 27, 90) [subsumption resolution 3459,2853] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3461. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 25, 26, 27, 28, 90) [subsumption resolution 3460,2858] {a:1,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3462. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 25, 26, 27, 28, 29, 90) [subsumption resolution 3461,2863] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3463. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~5, 25, 26, 27, 28, 29, 30, 90) [subsumption resolution 3462,2868] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3464. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) <- (~5, 25, 26, 27, 28, 29, 30, ~31, 90) [subsumption resolution 3463,2873] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3465. ~18 | ~17 | 5 | ~25 | ~26 | ~27 | ~28 | ~29 | ~30 | 31 | ~90 [avatar split clause 3464,3316,2871,2866,2861,2856,2851,2846,2841,1559,1687,1771]
3600. k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,sK5) | ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6) [superposition 1565,1541] {a:1,w:33,wCS:3300,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3601. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6) [trivial inequality removal 3600] {a:1,w:24,wCS:2400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3603. ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17) [subsumption resolution 3601,1689] {a:1,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3604. ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18) [subsumption resolution 3603,1773] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3605. ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 25) [subsumption resolution 3604,2843] {a:1,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3606. ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 25, 90) [subsumption resolution 3605,3318] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3607. ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 25, 26, 90) [subsumption resolution 3606,2848] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3608. ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 25, 26, 27, 90) [subsumption resolution 3607,2853] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3609. ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 25, 26, 27, 28, 90) [subsumption resolution 3608,2858] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3610. ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 25, 26, 27, 28, 29, 90) [subsumption resolution 3609,2863] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3611. v2_struct_0(sK0) <- (~6, 17, 18, 25, 26, 27, 28, 29, 30, 90) [subsumption resolution 3610,2868] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3612. $false <- (~6, 17, 18, 25, 26, 27, 28, 29, 30, ~31, 90) [subsumption resolution 3611,2873] {a:1,w:0,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3613. 6 | ~17 | ~18 | ~25 | ~26 | ~27 | ~28 | ~29 | ~30 | 31 | ~90 [avatar contradiction clause 3612]
3659. ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~1, 2) [subsumption resolution 3369,1545] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3660. ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 25) [subsumption resolution 3659,2843] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3661. ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 25, 90) [subsumption resolution 3660,3318] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3662. ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 25, 26, 90) [subsumption resolution 3661,2848] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3663. ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 25, 26, 27, 90) [subsumption resolution 3662,2853] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3664. ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 25, 26, 27, 28, 90) [subsumption resolution 3663,2858] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3665. ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~1, 2, 25, 26, 27, 28, 29, 90) [subsumption resolution 3664,2863] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3666. ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~1, 2, 25, 26, 27, 28, 29, 30, 90) [subsumption resolution 3665,2868] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3667. ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) <- (~1, 2, 25, 26, 27, 28, 29, 30, ~31, 90) [subsumption resolution 3666,2873] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3668. ~9 | ~8 | 1 | ~2 | ~25 | ~26 | ~27 | ~28 | ~29 | ~30 | 31 | ~90 [avatar split clause 3667,3316,2871,2866,2861,2856,2851,2846,2841,1547,1543,1573,1578]
3756. sK11 != sK11 | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1) [superposition 1537,1544] {a:1,w:37,wCS:3700,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3762. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1) [trivial inequality removal 3756] {a:1,w:34,wCS:3400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3791. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8) [subsumption resolution 3762,1575] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3792. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9) [subsumption resolution 3791,1580] {a:1,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3793. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 25) [subsumption resolution 3792,2843] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3794. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 25, 90) [subsumption resolution 3793,3318] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3795. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 25, 26, 90) [subsumption resolution 3794,2848] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3796. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 25, 26, 27, 90) [subsumption resolution 3795,2853] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3797. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 25, 26, 27, 28, 90) [subsumption resolution 3796,2858] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3798. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 25, 26, 27, 28, 29, 90) [subsumption resolution 3797,2863] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3799. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | v2_struct_0(sK0) <- (1, 8, 9, 25, 26, 27, 28, 29, 30, 90) [subsumption resolution 3798,2868] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3800. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) <- (1, 8, 9, 25, 26, 27, 28, 29, 30, ~31, 90) [subsumption resolution 3799,2873] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3801. 2 | ~1 | ~8 | ~9 | ~25 | ~26 | ~27 | ~28 | ~29 | ~30 | 31 | ~90 [avatar split clause 3800,3316,2871,2866,2861,2856,2851,2846,2841,1578,1573,1543,1547]
3897. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | ~r1_bcialg_1(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7) [superposition 1569,1536] {a:1,w:53,wCS:5300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3898. ~r1_bcialg_1(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7) [trivial inequality removal 3897] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3899. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19) [subsumption resolution 3898,1857] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3900. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20) [subsumption resolution 3899,2101] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3901. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 22) [subsumption resolution 3900,2589] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3999. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10) [subsumption resolution 3338,1553] {a:1,w:24,wCS:2400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4000. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 12) [subsumption resolution 3999,1601] {a:1,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4001. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 12, 25) [subsumption resolution 4000,2843] {a:1,w:18,wCS:1800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4002. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 12, 25, 90) [subsumption resolution 4001,3318] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4003. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 12, 25, 26, 90) [subsumption resolution 4002,2848] {a:1,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4004. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 12, 25, 26, 27, 90) [subsumption resolution 4003,2853] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4005. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 12, 25, 26, 27, 28, 90) [subsumption resolution 4004,2858] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4006. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 12, 25, 26, 27, 28, 29, 90) [subsumption resolution 4005,2863] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4007. ~m1_subset_1(sK10,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~3, 10, 12, 25, 26, 27, 28, 29, 30, 90) [subsumption resolution 4006,2868] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4008. ~m1_subset_1(sK10,u1_struct_0(sK0)) <- (~3, 10, 12, 25, 26, 27, 28, 29, 30, ~31, 90) [subsumption resolution 4007,2873] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4009. ~11 | 3 | ~10 | ~12 | ~25 | ~26 | ~27 | ~28 | ~29 | ~30 | 31 | ~90 [avatar split clause 4008,3316,2871,2866,2861,2856,2851,2846,2841,1599,1583,1551,1591]
4077. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 22, 25) [subsumption resolution 3901,2843] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4078. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 22, 25, 90) [subsumption resolution 4077,3318] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4079. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 22, 25, 26, 90) [subsumption resolution 4078,2848] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4080. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 22, 25, 26, 27, 90) [subsumption resolution 4079,2853] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4081. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 22, 25, 26, 27, 28, 90) [subsumption resolution 4080,2858] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4082. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 22, 25, 26, 27, 28, 29, 90) [subsumption resolution 4081,2863] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4083. ~m1_subset_1(sK2,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~7, 19, 20, 22, 25, 26, 27, 28, 29, 30, 90) [subsumption resolution 4082,2868] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4084. ~m1_subset_1(sK2,u1_struct_0(sK0)) <- (~7, 19, 20, 22, 25, 26, 27, 28, 29, 30, ~31, 90) [subsumption resolution 4083,2873] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4085. ~21 | 7 | ~19 | ~20 | ~22 | ~25 | ~26 | ~27 | ~28 | ~29 | ~30 | 31 | ~90 [avatar split clause 4084,3316,2871,2866,2861,2856,2851,2846,2841,2587,2099,1855,1567,2343]
4086. $false [avatar sat refutation 1570,1571,1576,1581,1586,1587,1588,1589,1594,1595,1596,1597,1602,1603,1604,1605,1610,1611,1612,1613,1614,1615,1616,1617,1618,1619,1620,1621,1622,1623,1624,1625,1630,1631,1632,1633,1634,1635,1636,1637,1638,1639,1640,1641,1642,1643,1644,1645,1650,1651,1652,1653,1654,1655,1656,1657,1658,1659,1660,1661,1662,1663,1664,1665,1670,1671,1672,1673,1674,1675,1676,1677,1678,1679,1680,1681,1682,1683,1684,1685,1690,1691,1692,1693,1694,1695,1696,1697,1698,1699,1700,1701,1702,1703,1704,1705,1706,1707,1708,1709,1710,1711,1712,1713,1714,1715,1716,1717,1718,1719,1720,1721,1722,1723,1724,1725,1726,1727,1728,1729,1730,1731,1732,1733,1734,1735,1736,1737,1738,1739,1740,1741,1742,1743,1744,1745,1746,1747,1748,1749,1750,1751,1752,1753,1754,1755,1756,1757,1758,1759,1760,1761,1762,1763,1764,1765,1766,1767,1768,1769,1774,1775,1776,1777,1778,1779,1780,1781,1782,1783,1784,1785,1786,1787,1788,1789,1790,1791,1792,1793,1794,1795,1796,1797,1798,1799,1800,1801,1802,1803,1804,1805,1806,1807,1808,1809,1810,1811,1812,1813,1814,1815,1816,1817,1818,1819,1820,1821,1822,1823,1824,1825,1826,1827,1828,1829,1830,1831,1832,1833,1834,1835,1836,1837,1838,1839,1840,1841,1842,1843,1844,1845,1846,1847,1848,1849,1850,1851,1852,1853,1858,1859,1860,1861,1862,1863,1864,1865,1866,1867,1868,1869,1870,1871,1872,1873,1874,1875,1876,1877,1878,1879,1880,1881,1882,1883,1884,1885,1886,1887,1888,1889,1890,1891,1892,1893,1894,1895,1896,1897,1898,1899,1900,1901,1902,1903,1904,1905,1906,1907,1908,1909,1910,1911,1912,1913,1914,1915,1916,1917,1918,1919,1920,1921,1922,1923,1924,1925,1926,1927,1928,1929,1930,1931,1932,1933,1934,1935,1936,1937,1938,1939,1940,1941,1942,1943,1944,1945,1946,1947,1948,1949,1950,1951,1952,1953,1954,1955,1956,1957,1958,1959,1960,1961,1962,1963,1964,1965,1966,1967,1968,1969,1970,1971,1972,1973,1974,1975,1976,1977,1978,1979,1980,1981,1982,1983,1984,1985,1986,1987,1988,1989,1990,1991,1992,1993,1994,1995,1996,1997,1998,1999,2000,2001,2002,2003,2004,2005,2006,2007,2008,2009,2010,2011,2012,2013,2014,2015,2016,2017,2018,2019,2020,2021,2022,2023,2024,2025,2026,2027,2028,2029,2030,2031,2032,2033,2034,2035,2036,2037,2038,2039,2040,2041,2042,2043,2044,2045,2046,2047,2048,2049,2050,2051,2052,2053,2054,2055,2056,2057,2058,2059,2060,2061,2062,2063,2064,2065,2066,2067,2068,2069,2070,2071,2072,2073,2074,2075,2076,2077,2078,2079,2080,2081,2082,2083,2084,2085,2086,2087,2088,2089,2090,2091,2092,2093,2094,2095,2096,2097,2102,2103,2104,2105,2106,2107,2108,2109,2110,2111,2112,2113,2114,2115,2116,2117,2118,2119,2120,2121,2122,2123,2124,2125,2126,2127,2128,2129,2130,2131,2132,2133,2134,2135,2136,2137,2138,2139,2140,2141,2142,2143,2144,2145,2146,2147,2148,2149,2150,2151,2152,2153,2154,2155,2156,2157,2158,2159,2160,2161,2162,2163,2164,2165,2166,2167,2168,2169,2170,2171,2172,2173,2174,2175,2176,2177,2178,2179,2180,2181,2182,2183,2184,2185,2186,2187,2188,2189,2190,2191,2192,2193,2194,2195,2196,2197,2198,2199,2200,2201,2202,2203,2204,2205,2206,2207,2208,2209,2210,2211,2212,2213,2214,2215,2216,2217,2218,2219,2220,2221,2222,2223,2224,2225,2226,2227,2228,2229,2230,2231,2232,2233,2234,2235,2236,2237,2238,2239,2240,2241,2242,2243,2244,2245,2246,2247,2248,2249,2250,2251,2252,2253,2254,2255,2256,2257,2258,2259,2260,2261,2262,2263,2264,2265,2266,2267,2268,2269,2270,2271,2272,2273,2274,2275,2276,2277,2278,2279,2280,2281,2282,2283,2284,2285,2286,2287,2288,2289,2290,2291,2292,2293,2294,2295,2296,2297,2298,2299,2300,2301,2302,2303,2304,2305,2306,2307,2308,2309,2310,2311,2312,2313,2314,2315,2316,2317,2318,2319,2320,2321,2322,2323,2324,2325,2326,2327,2328,2329,2330,2331,2332,2333,2334,2335,2336,2337,2338,2339,2340,2341,2346,2347,2348,2349,2350,2351,2352,2353,2354,2355,2356,2357,2358,2359,2360,2361,2362,2363,2364,2365,2366,2367,2368,2369,2370,2371,2372,2373,2374,2375,2376,2377,2378,2379,2380,2381,2382,2383,2384,2385,2386,2387,2388,2389,2390,2391,2392,2393,2394,2395,2396,2397,2398,2399,2400,2401,2402,2403,2404,2405,2406,2407,2408,2409,2410,2411,2412,2413,2414,2415,2416,2417,2418,2419,2420,2421,2422,2423,2424,2425,2426,2427,2428,2429,2430,2431,2432,2433,2434,2435,2436,2437,2438,2439,2440,2441,2442,2443,2444,2445,2446,2447,2448,2449,2450,2451,2452,2453,2454,2455,2456,2457,2458,2459,2460,2461,2462,2463,2464,2465,2466,2467,2468,2469,2470,2471,2472,2473,2474,2475,2476,2477,2478,2479,2480,2481,2482,2483,2484,2485,2486,2487,2488,2489,2490,2491,2492,2493,2494,2495,2496,2497,2498,2499,2500,2501,2502,2503,2504,2505,2506,2507,2508,2509,2510,2511,2512,2513,2514,2515,2516,2517,2518,2519,2520,2521,2522,2523,2524,2525,2526,2527,2528,2529,2530,2531,2532,2533,2534,2535,2536,2537,2538,2539,2540,2541,2542,2543,2544,2545,2546,2547,2548,2549,2550,2551,2552,2553,2554,2555,2556,2557,2558,2559,2560,2561,2562,2563,2564,2565,2566,2567,2568,2569,2570,2571,2572,2573,2574,2575,2576,2577,2578,2579,2580,2581,2582,2583,2584,2585,2590,2591,2592,2593,2594,2595,2596,2597,2598,2599,2600,2601,2602,2603,2604,2605,2606,2607,2608,2609,2610,2611,2612,2613,2614,2615,2616,2617,2618,2619,2620,2621,2622,2623,2624,2625,2626,2627,2628,2629,2630,2631,2632,2633,2634,2635,2636,2637,2638,2639,2640,2641,2642,2643,2644,2645,2646,2647,2648,2649,2650,2651,2652,2653,2654,2655,2656,2657,2658,2659,2660,2661,2662,2663,2664,2665,2666,2667,2668,2669,2670,2671,2672,2673,2674,2675,2676,2677,2678,2679,2680,2681,2682,2683,2684,2685,2686,2687,2688,2689,2690,2691,2692,2693,2694,2695,2696,2697,2698,2699,2700,2701,2702,2703,2704,2705,2706,2707,2708,2709,2710,2711,2712,2713,2714,2715,2716,2717,2718,2719,2720,2721,2722,2723,2724,2725,2726,2727,2728,2729,2730,2731,2732,2733,2734,2735,2736,2737,2738,2739,2740,2741,2742,2743,2744,2745,2746,2747,2748,2749,2750,2751,2752,2753,2754,2755,2756,2757,2758,2759,2760,2761,2762,2763,2764,2765,2766,2767,2768,2769,2770,2771,2772,2773,2774,2775,2776,2777,2778,2779,2780,2781,2782,2783,2784,2785,2786,2787,2788,2789,2790,2791,2792,2793,2794,2795,2796,2797,2798,2799,2800,2801,2802,2803,2804,2805,2806,2807,2808,2809,2810,2811,2812,2813,2814,2815,2816,2817,2818,2819,2820,2821,2822,2823,2824,2825,2826,2827,2828,2829,2834,2839,2844,2849,2854,2859,2864,2869,2874,3319,3455,3465,3613,3668,3801,4009,4085] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:26346,thDist:-26346}
% SZS output end Proof for bciideal__t44_bciideal
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 73
% >>> Preprocessing
% Introduced skolems: 28
% Pure predicates: 3
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 1306
% Generated clauses: 2182
% Active clauses: 181
% Passive clauses: 260
% Final active clauses: 144
% Final passive clauses: 70
% Main loop iterations started: 186
% 
% >>> Simplifying Inferences
% Duplicate literals: 330
% Trivial inequalities: 12
% Fw subsumption resolutions: 428
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 27
% Forward subsumptions: 60
% 
% >>> Generating Inferences
% Binary resolution: 135
% Unit resulting resolution: 3
% Forward superposition: 33
% Backward superposition: 35
% Self superposition: 6
% 
% >>> AVATAR
% Split clauses: 1252
% Split components: 7134
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 1335
% SAT solver unit clauses: 27
% SAT solver binary clauses: 24
% 
% Memory used [KB]: 14711
% Time elapsed: 0.038 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% preprocessing: 0.001 s
% SAT solver time: 0.002 s
% sat proof minimization: 0.007 s
% splitting component index usage: 0.001 s
% forward subsumption: 0.002 s
% forward subsumption resolution: 0.001 s
% forward demodulation: 0.002 s
% unit resulting resolution: 0.001 s
% superposition: 0.001 s
% other: 0.038 s ( own 0.019 s ) 

