i: [1,69,0,3,1,1,0,0] l2_subset_1
a: [3,105,0,3,1,69]
s: 105
i: [1,68,0,3,1,1,0,0] l2_subset_1
a: [3,110,0,3,1,68]
s: 110
i: [1,67,0,4,1,1,0,0] l2_subset_1
a: [3,115,0,4,1,67]
s: 115
i: [1,70,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,120,0,2,1,70]
s: 120
i: [1,71,0,2,1,0,0,255] fc1_xboole_0
a: [3,125,0,2,1,71]
s: 125
i: [1,95,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,135,0,2,1,95]
s: 135
i: [1,96,0,2,1,0,0,2] rc2_xboole_0
a: [3,140,0,2,1,96]
s: 140
i: [1,97,0,2,1,0,0,2] rc1_xboole_0
a: [3,145,0,2,1,97]
s: 145
i: [1,99,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,150,0,2,1,99]
s: 150
i: [1,98,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,155,0,2,1,98]
s: 155
i: [1,73,0,3,1,0,0,1] fc1_subset_1
s: 73
i: [1,93,0,3,1,0,0,255] reflexivity_r1_tarski
s: 93
i: [1,72,0,3,1,0,0,255] d2_xboole_0
a: [3,130,0,3,1,72]
s: 130
r: 120
i: [1,74,0,4,2,0,0,255] cc2_zfmisc_1
s: 74
i: [1,77,0,4,2,0,0,255] cc1_zfmisc_1
s: 77
i: [1,92,0,4,1,0,0,1] existence_m1_subset_1
s: 92
d: [2,158,1,2,1,1,40,74,145]
a: [3,163,1,2,1,158]
s: 163
i: [1,76,0,5,2,0,0,1] rc1_subset_1
s: 76
d: [2,159,1,2,1,1,40,74,125]
a: [3,168,1,2,1,159]
s: 168
i: [1,78,0,5,2,0,0,255] t6_boole
s: 78
i: [1,91,0,5,2,0,0,1] t7_boole
s: 91
d: [2,173,1,2,1,1,40,91,110]
a: [3,177,1,2,1,173]
s: 177
i: [1,94,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 94
d: [2,179,1,3,1,1,40,94,110]
a: [3,183,1,3,1,179]
s: 183
i: [1,75,0,7,2,0,0,1] rc1_subset_1
s: 75
i: [1,81,0,7,3,0,0,1] d1_subset_1
s: 81
d: [2,186,1,5,2,0,40,81,92]
s: 186
i: [1,82,0,7,3,0,0,1] d1_subset_1
s: 82
d: [2,188,2,5,2,0,40,186,74]
s: 188
i: [1,90,0,7,3,0,0,255] t8_boole
s: 90
i: [1,100,0,7,2,0,0,255] d1_zfmisc_1
s: 100
i: [1,101,0,7,2,0,0,255] d1_zfmisc_1
s: 101
i: [1,79,0,8,3,0,0,1] d1_subset_1
s: 79
d: [2,195,1,7,2,1,40,79,115]
d: [2,198,1,4,1,1,21,195,73]
a: [3,202,1,4,1,198]
s: 202
i: [1,80,0,8,3,0,0,1] d1_subset_1
s: 80
d: [2,218,0,6,2,0,21,80,91]
s: 218
d: [2,219,1,3,1,1,40,218,110]
a: [3,225,1,3,1,219]
s: 225
i: [1,84,0,8,2,0,0,255] d3_tarski
s: 84
d: [2,205,2,3,1,1,40,202,101]
a: [3,211,2,3,1,205]
s: 211
i: [1,85,0,8,2,0,0,255] d3_tarski
s: 85
i: [1,83,0,9,3,0,0,255] d3_tarski
s: 83
d: [2,236,3,0,0,3,52,105,110,211,83]
e: 236
f: 236,205,198,67,68,69
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__l2_subset_1
% SZS output start Proof for subset_1__l2_subset_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (r2_hidden(X2,X1) => r2_hidden(X2,X0))) [input l2_subset_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (r2_hidden(X2,X1) => r2_hidden(X2,X0))) [negated conjecture 1]
6. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
7. ! [X0] : ! [X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [input d1_zfmisc_1]
9. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
15. ! [X0] : ~v1_xboole_0(k1_zfmisc_1(X0)) [input fc1_subset_1]
26. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => ! [X2] : (r2_hidden(X2,X1) => r2_hidden(X2,X0))) [flattening 2]
27. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 6]
28. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 9]
29. ! [X0,X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [flattening 7]
33. ? [X0,X1] : (? [X2] : (~r2_hidden(X2,X0) & r2_hidden(X2,X1)) & m1_subset_1(X1,k1_zfmisc_1(X0))) [ennf transformation 26]
38. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 27]
39. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 28]
43. ? [X0,X1] : (? [X2] : (~r2_hidden(X2,X0) & r2_hidden(X2,X1)) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (? [X2] : (~r2_hidden(X2,sK0) & r2_hidden(X2,sK1)) & m1_subset_1(sK1,k1_zfmisc_1(sK0))) [choice axiom]
44. ? [X2] : (~r2_hidden(X2,sK0) & r2_hidden(X2,sK1)) => (~r2_hidden(sK2,sK0) & r2_hidden(sK2,sK1)) [choice axiom]
45. (~r2_hidden(sK2,sK0) & r2_hidden(sK2,sK1)) & m1_subset_1(sK1,k1_zfmisc_1(sK0)) [skolemisation 33,44,43]
48. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 38]
49. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 39]
50. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 49]
51. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK4(X0,X1),X0))) [choice axiom]
52. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK4(X0,X1),X1) & r2_hidden(sK4(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 50,51]
53. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ~r1_tarski(X2,X0)) & (r1_tarski(X2,X0) | ~r2_hidden(X2,X1))) | k1_zfmisc_1(X0) != X1)) [nnf transformation 29]
54. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [rectify 53]
55. ! [X1,X0] : (? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1))) => ((~r1_tarski(sK5(X0,X1),X0) | ~r2_hidden(sK5(X0,X1),X1)) & (r1_tarski(sK5(X0,X1),X0) | r2_hidden(sK5(X0,X1),X1)))) [choice axiom]
56. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ((~r1_tarski(sK5(X0,X1),X0) | ~r2_hidden(sK5(X0,X1),X1)) & (r1_tarski(sK5(X0,X1),X0) | r2_hidden(sK5(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [skolemisation 54,55]
67. m1_subset_1(sK1,k1_zfmisc_1(sK0)) [cnf transformation 45] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
68. r2_hidden(sK2,sK1) [cnf transformation 45] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
69. ~r2_hidden(sK2,sK0) [cnf transformation 45] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
73. ~v1_xboole_0(k1_zfmisc_1(X0)) [cnf transformation 15] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. ~m1_subset_1(X1,X0) | r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 48] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
83. ~r1_tarski(X0,X1) | ~r2_hidden(X3,X0) | r2_hidden(X3,X1) [cnf transformation 52] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
86. r1_tarski(X3,X0) | ~r2_hidden(X3,X1) | k1_zfmisc_1(X0) != X1 [cnf transformation 56] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
101. ~r2_hidden(X3,k1_zfmisc_1(X0)) | r1_tarski(X3,X0) [equality resolution 86] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
103. 1 <=> r2_hidden(sK2,sK0) [avatar definition]
105. ~r2_hidden(sK2,sK0) <- (~1) [avatar component clause 103] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
106. ~1 [avatar split clause 69,103]
108. 2 <=> r2_hidden(sK2,sK1) [avatar definition]
110. r2_hidden(sK2,sK1) <- (2) [avatar component clause 108] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
111. 2 [avatar split clause 68,108]
113. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(sK0)) [avatar definition]
115. m1_subset_1(sK1,k1_zfmisc_1(sK0)) <- (3) [avatar component clause 113] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. 3 [avatar split clause 67,113]
195. r2_hidden(sK1,k1_zfmisc_1(sK0)) | v1_xboole_0(k1_zfmisc_1(sK0)) <- (3) [resolution 79,115] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
198. r2_hidden(sK1,k1_zfmisc_1(sK0)) <- (3) [subsumption resolution 195,73] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
200. 16 <=> r2_hidden(sK1,k1_zfmisc_1(sK0)) [avatar definition]
202. r2_hidden(sK1,k1_zfmisc_1(sK0)) <- (16) [avatar component clause 200] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
203. 16 | ~3 [avatar split clause 198,113,200]
205. r1_tarski(sK1,sK0) <- (16) [resolution 202,101] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
209. 17 <=> r1_tarski(sK1,sK0) [avatar definition]
211. r1_tarski(sK1,sK0) <- (17) [avatar component clause 209] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
212. 17 | ~16 [avatar split clause 205,200,209]
236. $false <- (~1, 2, 17) [unit resulting resolution 105,110,211,83] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
239. 1 | ~2 | ~17 [avatar contradiction clause 236]
240. $false [avatar sat refutation 106,111,116,203,212,239] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:95,thDist:-95}
% SZS output end Proof for subset_1__l2_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% Introduced skolems: 11
% 
% >>> Saturation
% Initial clauses: 33
% Generated clauses: 94
% Active clauses: 39
% Passive clauses: 54
% Final active clauses: 38
% Final passive clauses: 13
% Main loop iterations started: 41
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Fw subsumption resolutions: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 32
% Unit resulting resolution: 1
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 22
% SAT solver unit clauses: 11
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

