i: [1,90,0,2,1,1,0,0] t88_mcart_1
a: [3,142,0,2,1,90]
s: 142
i: [1,91,0,3,1,1,0,0] t88_mcart_1
a: [3,137,0,3,1,91]
s: 137
i: [1,92,0,7,1,1,0,0] t88_mcart_1
a: [3,132,0,7,1,92]
s: 132
i: [1,93,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,147,0,2,1,93]
s: 147
i: [1,94,0,2,1,0,0,255] fc1_xboole_0
a: [3,152,0,2,1,94]
s: 152
i: [1,121,0,2,1,0,0,2] rc2_xboole_0
a: [3,162,0,2,1,121]
s: 162
i: [1,122,0,2,1,0,0,2] rc1_xboole_0
a: [3,167,0,2,1,122]
s: 167
i: [1,124,0,2,1,0,0,1] rc1_relat_1
a: [3,172,0,2,1,124]
s: 172
i: [1,123,0,2,1,0,0,1] rc1_relat_1
a: [3,177,0,2,1,123]
s: 177
i: [1,125,0,2,1,0,0,255] rc2_relat_1
a: [3,182,0,2,1,125]
s: 182
i: [1,96,0,3,1,0,0,255] fc2_xboole_0
s: 96
i: [1,95,0,3,1,0,0,255] d2_xboole_0
a: [3,157,0,3,1,95]
s: 157
r: 147
i: [1,98,0,4,2,0,0,1] cc1_relat_1
s: 98
d: [2,185,1,2,1,1,40,98,167]
a: [3,190,1,2,1,185]
s: 190
i: [1,109,0,4,1,0,0,255] existence_m1_subset_1
s: 109
i: [1,99,0,5,2,0,0,255] t6_boole
s: 99
d: [2,186,1,2,1,1,40,98,152]
a: [3,195,1,2,1,186]
s: 195
i: [1,106,0,5,2,0,0,1] t7_boole
s: 106
d: [2,200,1,2,1,1,40,106,137]
a: [3,204,1,2,1,200]
s: 204
i: [1,104,0,6,2,0,0,255] t1_subset
s: 104
i: [1,110,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 110
d: [2,206,1,3,1,1,40,104,137]
a: [3,210,1,3,1,206]
s: 210
i: [1,105,0,7,3,0,0,255] t8_boole
s: 105
d: [2,212,1,3,1,1,40,110,137]
a: [3,216,1,3,1,212]
s: 216
i: [1,107,0,7,2,0,0,1] t2_tarski
s: 107
i: [1,103,0,8,3,0,0,255] t2_subset
s: 103
d: [2,221,1,6,2,0,40,103,109]
s: 221
i: [1,111,0,8,3,0,0,255] fc19_relat_1
s: 111
d: [2,223,2,6,2,0,40,221,110]
s: 223
i: [1,112,0,8,3,0,0,255] fc20_relat_1
s: 112
d: [2,227,0,6,2,0,21,112,98]
s: 227
d: [2,228,1,6,2,0,40,227,98]
s: 228
i: [1,97,0,10,2,0,0,255] d16_relat_1
s: 97
d: [2,230,1,8,3,0,44,227,97]
d: [2,234,1,6,2,0,21,230,98]
s: 234
i: [1,108,0,10,3,0,0,1] t2_tarski
s: 108
i: [1,100,0,11,3,0,0,255] t87_mcart_1
s: 100
d: [2,229,2,8,3,0,44,228,97]
d: [2,233,2,6,2,0,21,229,98]
s: 233
i: [1,113,0,13,3,0,0,255] fraenkel_a_2_1_mcart_1
s: 113
d: [2,235,2,6,2,0,40,234,98]
s: 235
i: [1,114,0,14,3,0,0,255] fraenkel_a_2_1_mcart_1
s: 114
i: [1,128,0,14,4,0,0,255] fraenkel_a_2_0_mcart_1
s: 128
d: [2,197,1,7,3,0,44,99,99]
s: 197
i: [1,115,0,15,3,0,0,255] fraenkel_a_2_1_mcart_1
s: 115
d: [2,218,1,7,2,0,40,107,110]
s: 218
i: [1,117,0,15,4,0,0,255] fraenkel_a_2_0_mcart_1
s: 117
i: [1,118,0,16,4,0,0,255] fraenkel_a_2_0_mcart_1
s: 118
d: [2,219,1,7,2,0,40,107,104]
s: 219
i: [1,119,0,16,4,0,0,255] fraenkel_a_2_0_mcart_1
s: 119
d: [2,226,1,8,3,0,40,111,98]
s: 226
i: [1,126,0,16,4,0,0,255] fraenkel_a_2_1_mcart_1
s: 126
d: [2,257,1,11,3,0,47,126]
s: 257
d: [2,260,2,9,3,0,40,257,106]
s: 260
d: [2,236,1,12,3,0,40,108,107]
s: 236
d: [2,237,2,11,3,0,40,108,221]
s: 237
d: [2,241,1,12,4,0,40,128,106]
s: 241
d: [2,239,1,14,4,0,40,128,110]
s: 239
d: [2,258,2,11,3,0,40,257,110]
s: 258
d: [2,240,1,14,4,0,40,128,104]
s: 240
d: [2,259,2,11,3,0,40,257,104]
s: 259
d: [2,243,1,18,6,0,44,128,100]
d: [2,244,1,14,4,0,18,243]
s: 244
d: [2,285,2,0,0,4,52,204,142,210,132,244]
e: 285
d: [2,286,2,7,3,1,40,244,132]
d: [2,295,2,4,2,2,21,286,210]
d: [2,296,2,2,1,3,21,295,142]
d: [2,297,2,0,0,4,21,296,204]
e: 297
f: 297,206,200,90,91,92
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t88_mcart_1
% SZS output start Proof for mcart_1__t88_mcart_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r2_hidden(X1,X0) => r2_hidden(k2_xtuple_0(X1),k9_relat_1(X0,k1_xtuple_0(X1))))) [input t88_mcart_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (r2_hidden(X1,X0) => r2_hidden(k2_xtuple_0(X1),k9_relat_1(X0,k1_xtuple_0(X1))))) [negated conjecture 1]
22. ! [X0,X1,X2] : ((v1_relat_1(X1) & ~v1_xboole_0(X1)) => (r2_hidden(X0,a_2_0_mcart_1(X1,X2)) <=> ? [X3] : (k1_xtuple_0(X3) = X2 & k2_xtuple_0(X3) = X0 & m1_subset_1(X3,X1)))) [input fraenkel_a_2_0_mcart_1]
28. ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [input t1_subset]
32. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
33. ! [X0] : ((v1_relat_1(X0) & ~v1_xboole_0(X0)) => ! [X1] : k9_relat_1(X0,X1) = a_2_0_mcart_1(X0,X1)) [input t87_mcart_1]
36. ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1)) [flattening 28]
38. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 32]
43. ? [X0] : (? [X1] : (~r2_hidden(k2_xtuple_0(X1),k9_relat_1(X0,k1_xtuple_0(X1))) & r2_hidden(X1,X0)) & v1_relat_1(X0)) [ennf transformation 2]
47. ! [X0] : (! [X1] : k9_relat_1(X0,X1) = a_2_0_mcart_1(X0,X1) | (~v1_relat_1(X0) | v1_xboole_0(X0))) [ennf transformation 33]
48. ! [X0] : (! [X1] : k9_relat_1(X0,X1) = a_2_0_mcart_1(X0,X1) | ~v1_relat_1(X0) | v1_xboole_0(X0)) [flattening 47]
55. ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1)) [ennf transformation 36]
57. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 38]
65. ! [X0,X1,X2] : ((r2_hidden(X0,a_2_0_mcart_1(X1,X2)) <=> ? [X3] : (k1_xtuple_0(X3) = X2 & k2_xtuple_0(X3) = X0 & m1_subset_1(X3,X1))) | (~v1_relat_1(X1) | v1_xboole_0(X1))) [ennf transformation 22]
66. ! [X0,X1,X2] : ((r2_hidden(X0,a_2_0_mcart_1(X1,X2)) <=> ? [X3] : (k1_xtuple_0(X3) = X2 & k2_xtuple_0(X3) = X0 & m1_subset_1(X3,X1))) | ~v1_relat_1(X1) | v1_xboole_0(X1)) [flattening 65]
67. ? [X0] : (? [X1] : (~r2_hidden(k2_xtuple_0(X1),k9_relat_1(X0,k1_xtuple_0(X1))) & r2_hidden(X1,X0)) & v1_relat_1(X0)) => (? [X1] : (~r2_hidden(k2_xtuple_0(X1),k9_relat_1(sK0,k1_xtuple_0(X1))) & r2_hidden(X1,sK0)) & v1_relat_1(sK0)) [choice axiom]
68. ? [X1] : (~r2_hidden(k2_xtuple_0(X1),k9_relat_1(sK0,k1_xtuple_0(X1))) & r2_hidden(X1,sK0)) => (~r2_hidden(k2_xtuple_0(sK1),k9_relat_1(sK0,k1_xtuple_0(sK1))) & r2_hidden(sK1,sK0)) [choice axiom]
69. (~r2_hidden(k2_xtuple_0(sK1),k9_relat_1(sK0,k1_xtuple_0(sK1))) & r2_hidden(sK1,sK0)) & v1_relat_1(sK0) [skolemisation 43,68,67]
78. ! [X0,X1,X2] : (((r2_hidden(X0,a_2_0_mcart_1(X1,X2)) | ! [X3] : (k1_xtuple_0(X3) != X2 | k2_xtuple_0(X3) != X0 | ~m1_subset_1(X3,X1))) & (? [X3] : (k1_xtuple_0(X3) = X2 & k2_xtuple_0(X3) = X0 & m1_subset_1(X3,X1)) | ~r2_hidden(X0,a_2_0_mcart_1(X1,X2)))) | ~v1_relat_1(X1) | v1_xboole_0(X1)) [nnf transformation 66]
79. ! [X0,X1,X2] : (((r2_hidden(X0,a_2_0_mcart_1(X1,X2)) | ! [X3] : (k1_xtuple_0(X3) != X2 | k2_xtuple_0(X3) != X0 | ~m1_subset_1(X3,X1))) & (? [X4] : (k1_xtuple_0(X4) = X2 & k2_xtuple_0(X4) = X0 & m1_subset_1(X4,X1)) | ~r2_hidden(X0,a_2_0_mcart_1(X1,X2)))) | ~v1_relat_1(X1) | v1_xboole_0(X1)) [rectify 78]
80. ! [X2,X1,X0] : (? [X4] : (k1_xtuple_0(X4) = X2 & k2_xtuple_0(X4) = X0 & m1_subset_1(X4,X1)) => (k1_xtuple_0(sK5(X0,X1,X2)) = X2 & k2_xtuple_0(sK5(X0,X1,X2)) = X0 & m1_subset_1(sK5(X0,X1,X2),X1))) [choice axiom]
81. ! [X0,X1,X2] : (((r2_hidden(X0,a_2_0_mcart_1(X1,X2)) | ! [X3] : (k1_xtuple_0(X3) != X2 | k2_xtuple_0(X3) != X0 | ~m1_subset_1(X3,X1))) & ((k1_xtuple_0(sK5(X0,X1,X2)) = X2 & k2_xtuple_0(sK5(X0,X1,X2)) = X0 & m1_subset_1(sK5(X0,X1,X2),X1)) | ~r2_hidden(X0,a_2_0_mcart_1(X1,X2)))) | ~v1_relat_1(X1) | v1_xboole_0(X1)) [skolemisation 79,80]
90. v1_relat_1(sK0) [cnf transformation 69] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
91. r2_hidden(sK1,sK0) [cnf transformation 69] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
92. ~r2_hidden(k2_xtuple_0(sK1),k9_relat_1(sK0,k1_xtuple_0(sK1))) [cnf transformation 69] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. k9_relat_1(X0,X1) = a_2_0_mcart_1(X0,X1) | ~v1_relat_1(X0) | v1_xboole_0(X0) [cnf transformation 48] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
104. ~r2_hidden(X0,X1) | m1_subset_1(X0,X1) [cnf transformation 55] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
106. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 57] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. r2_hidden(X0,a_2_0_mcart_1(X1,X2)) | k1_xtuple_0(X3) != X2 | k2_xtuple_0(X3) != X0 | ~m1_subset_1(X3,X1) | ~v1_relat_1(X1) | v1_xboole_0(X1) [cnf transformation 81] {a:0,w:20,wCS:8000,sine:255,thAx:0,allAx:2,thDist:-2}
127. r2_hidden(X0,a_2_0_mcart_1(X1,k1_xtuple_0(X3))) | k2_xtuple_0(X3) != X0 | ~m1_subset_1(X3,X1) | ~v1_relat_1(X1) | v1_xboole_0(X1) [equality resolution 120] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:2,thDist:-2}
128. r2_hidden(k2_xtuple_0(X3),a_2_0_mcart_1(X1,k1_xtuple_0(X3))) | ~m1_subset_1(X3,X1) | ~v1_relat_1(X1) | v1_xboole_0(X1) [equality resolution 127] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
130. 1 <=> r2_hidden(k2_xtuple_0(sK1),k9_relat_1(sK0,k1_xtuple_0(sK1))) [avatar definition]
132. ~r2_hidden(k2_xtuple_0(sK1),k9_relat_1(sK0,k1_xtuple_0(sK1))) <- (~1) [avatar component clause 130] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
133. ~1 [avatar split clause 92,130]
135. 2 <=> r2_hidden(sK1,sK0) [avatar definition]
137. r2_hidden(sK1,sK0) <- (2) [avatar component clause 135] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
138. 2 [avatar split clause 91,135]
140. 3 <=> v1_relat_1(sK0) [avatar definition]
142. v1_relat_1(sK0) <- (3) [avatar component clause 140] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
143. 3 [avatar split clause 90,140]
200. ~v1_xboole_0(sK0) <- (2) [resolution 106,137] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
202. 14 <=> v1_xboole_0(sK0) [avatar definition]
204. ~v1_xboole_0(sK0) <- (~14) [avatar component clause 202] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
205. ~14 | ~2 [avatar split clause 200,135,202]
206. m1_subset_1(sK1,sK0) <- (2) [resolution 104,137] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
208. 15 <=> m1_subset_1(sK1,sK0) [avatar definition]
210. m1_subset_1(sK1,sK0) <- (15) [avatar component clause 208] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
211. 15 | ~2 [avatar split clause 206,135,208]
243. r2_hidden(k2_xtuple_0(X1),k9_relat_1(X0,k1_xtuple_0(X1))) | ~m1_subset_1(X1,X0) | ~v1_relat_1(X0) | v1_xboole_0(X0) | ~v1_relat_1(X0) | v1_xboole_0(X0) [superposition 128,100] {a:1,w:18,wCS:7200,sine:255,thAx:0,allAx:3,thDist:-3}
244. r2_hidden(k2_xtuple_0(X1),k9_relat_1(X0,k1_xtuple_0(X1))) | ~m1_subset_1(X1,X0) | ~v1_relat_1(X0) | v1_xboole_0(X0) [duplicate literal removal 243] {a:1,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
286. ~m1_subset_1(sK1,sK0) | ~v1_relat_1(sK0) | v1_xboole_0(sK0) <- (~1) [resolution 244,132] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
295. ~v1_relat_1(sK0) | v1_xboole_0(sK0) <- (~1, 15) [subsumption resolution 286,210] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
296. v1_xboole_0(sK0) <- (~1, 3, 15) [subsumption resolution 295,142] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
297. $false <- (~1, 3, ~14, 15) [subsumption resolution 296,204] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
298. 1 | ~3 | 14 | ~15 [avatar contradiction clause 297]
299. $false [avatar sat refutation 133,138,143,205,211,298] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:83,thDist:-83}
% SZS output end Proof for mcart_1__t88_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 33
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 140
% Active clauses: 57
% Passive clauses: 89
% Final active clauses: 56
% Final passive clauses: 29
% Main loop iterations started: 60
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 33
% Unit resulting resolution: 1
% Forward superposition: 31
% Backward superposition: 7
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 11
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

