i: [1,47,0,2,1,1,0,0] t68_relat_1
a: [3,79,0,2,1,47]
s: 79
i: [1,48,0,4,1,1,0,0] t68_relat_1
a: [3,74,0,4,1,48]
s: 74
i: [1,49,0,5,1,1,0,0] t68_relat_1
a: [3,69,0,5,1,49]
s: 69
i: [1,60,0,3,1,0,0,1] reflexivity_r1_tarski
s: 60
i: [1,59,0,4,1,0,0,255] existence_m1_subset_1
s: 59
i: [1,61,0,4,1,0,0,2] fc6_relat_1
s: 61
i: [1,62,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 62
i: [1,64,0,6,2,0,0,1] dt_k5_relat_1
s: 64
i: [1,65,0,6,2,0,0,2] fc1_relat_1
s: 65
i: [1,54,0,7,2,0,0,255] t3_subset
s: 54
i: [1,55,0,7,2,0,0,255] t3_subset
s: 55
d: [2,82,1,5,1,0,40,54,59]
s: 82
i: [1,63,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 63
d: [2,84,1,6,2,0,44,65,63]
s: 84
i: [1,51,0,8,3,0,0,255] cc2_relat_1
s: 51
i: [1,53,0,8,2,0,0,1] t28_xboole_1
s: 53
d: [2,85,1,6,2,0,44,65,63]
s: 85
i: [1,50,0,9,2,0,0,1] t7_relat_1
s: 50
d: [2,89,1,6,2,0,40,51,59]
s: 89
i: [1,58,0,9,3,0,0,1] t1_xboole_1
s: 58
i: [1,56,0,10,2,0,0,2] t95_zfmisc_1
s: 56
d: [2,97,1,7,2,1,40,58,74]
s: 97
i: [1,57,0,10,2,0,0,2] t95_zfmisc_1
s: 57
d: [2,103,2,6,1,1,40,97,82]
a: [3,107,2,6,1,103]
s: 107
i: [1,52,0,12,2,0,0,2] t67_relat_1
s: 52
d: [2,90,1,7,3,0,40,51,55]
s: 90
d: [2,118,2,5,2,1,40,90,74]
a: [3,130,2,2,1,118]
s: 130
d: [2,91,1,8,2,0,44,53,63]
s: 91
d: [2,110,3,9,2,1,40,107,58]
s: 110
d: [2,92,1,8,2,0,44,53,63]
s: 92
d: [2,94,1,8,2,0,44,63,53]
s: 94
d: [2,148,4,8,1,1,40,110,82]
a: [3,152,4,8,1,148]
s: 152
d: [2,95,1,8,2,0,44,63,53]
s: 95
d: [2,155,5,11,2,1,40,152,58]
s: 155
d: [2,99,1,12,3,0,40,58,50]
s: 99
d: [2,101,1,13,3,0,40,56,58]
s: 101
d: [2,157,6,10,1,1,40,155,82]
a: [3,161,6,10,1,157]
s: 161
d: [2,109,1,13,3,0,40,57,58]
s: 109
d: [2,172,7,13,2,1,40,161,58]
s: 172
d: [2,111,1,13,3,0,44,52,53]
s: 111
d: [2,112,1,13,3,0,44,53,52]
s: 112
d: [2,180,8,12,1,1,40,172,82]
a: [3,184,8,12,1,180]
s: 184
d: [2,93,2,7,3,0,44,84,53]
s: 93
d: [2,194,9,15,2,1,40,184,58]
s: 194
d: [2,100,2,8,2,0,40,58,82]
s: 100
d: [2,186,2,10,3,0,40,111,50]
d: [2,190,2,8,2,0,18,186]
s: 190
d: [2,196,10,14,1,1,40,194,82]
a: [3,200,10,14,1,196]
s: 200
d: [2,187,2,13,4,0,40,111,99]
d: [2,189,2,11,3,0,18,187]
d: [2,191,2,8,2,0,21,189,60]
s: 191
d: [2,207,11,17,2,1,40,200,58]
s: 207
d: [2,137,2,9,3,0,44,91,53]
s: 137
d: [2,141,2,9,3,0,44,53,91]
s: 141
d: [2,218,3,8,2,1,40,137,74]
a: [3,230,3,4,1,218]
s: 230
d: [2,165,2,12,3,0,40,101,50]
s: 165
d: [2,222,3,12,2,1,40,137,107]
a: [3,239,3,6,1,222]
s: 239
d: [2,170,2,12,2,0,40,101,82]
s: 170
d: [2,173,2,12,3,0,40,109,50]
s: 173
d: [2,223,5,16,2,1,40,137,152]
a: [3,248,5,8,1,223]
s: 248
d: [2,178,2,12,2,0,40,109,82]
s: 178
d: [2,224,7,20,2,1,40,137,161]
a: [3,257,7,10,1,224]
s: 257
d: [2,168,2,13,3,0,40,101,56]
s: 168
d: [2,169,2,13,3,0,40,101,57]
s: 169
d: [2,225,9,24,2,1,40,137,184]
a: [3,266,9,12,1,225]
s: 266
d: [2,176,2,13,3,0,40,109,56]
s: 176
d: [2,226,11,28,2,1,40,137,200]
a: [3,275,11,14,1,226]
s: 275
d: [2,177,2,13,3,0,40,109,57]
s: 177
d: [2,164,2,15,4,0,40,99,58]
s: 164
d: [2,209,12,16,1,1,40,207,82]
a: [3,213,12,16,1,209]
s: 213
d: [2,166,2,15,4,0,40,101,99]
s: 166
d: [2,335,13,32,2,1,40,213,137]
a: [3,341,13,16,1,335]
s: 341
d: [2,174,2,15,4,0,40,109,99]
s: 174
d: [2,188,2,19,3,0,40,111,57]
d: [2,192,2,15,2,0,21,188,61]
s: 192
d: [2,337,13,19,2,1,40,213,58]
s: 337
d: [2,140,2,16,3,0,44,91,52]
s: 140
d: [2,364,14,18,1,1,40,337,82]
a: [3,368,14,18,1,364]
s: 368
d: [2,146,2,16,3,0,44,52,91]
s: 146
d: [2,203,3,7,1,0,40,100,82]
s: 203
d: [2,404,15,36,2,1,40,368,137]
a: [3,410,15,18,1,404]
s: 410
d: [2,291,3,13,3,0,40,170,90]
d: [2,293,3,9,2,0,21,291,61]
s: 293
d: [2,427,14,20,1,1,40,203,337]
a: [3,432,14,20,1,427]
s: 432
d: [2,305,3,13,3,0,40,178,90]
d: [2,307,3,9,2,0,21,305,61]
s: 307
d: [2,221,3,10,2,0,40,137,82]
s: 221
d: [2,435,15,40,2,1,40,432,137]
a: [3,441,15,20,1,435]
s: 441
d: [2,331,3,16,5,0,40,164,90]
d: [2,334,3,10,4,0,21,331,61]
s: 334
d: [2,456,4,8,3,1,40,334,74]
a: [3,476,4,5,2,456]
s: 476
d: [2,281,3,13,4,0,40,165,111]
d: [2,287,3,11,3,0,18,281]
s: 287
d: [2,523,4,0,0,3,52,79,74,69,287]
e: 523
d: [2,524,4,9,3,1,44,69,287]
d: [2,527,4,6,2,1,19,524]
d: [2,529,4,4,1,2,21,527,79]
d: [2,530,4,0,0,3,21,529,74]
e: 530
f: 530,47,48,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t68_relat_1
% SZS output start Proof for relat_1__t68_relat_1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => (r1_tarski(k9_xtuple_0(X1),X0) => k5_relat_1(X1,X0) = X1)) [input t68_relat_1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => (r1_tarski(k9_xtuple_0(X1),X0) => k5_relat_1(X1,X0) = X1)) [negated conjecture 1]
17. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
18. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [input t28_xboole_1]
20. ! [X0] : ! [X1] : (v1_relat_1(X1) => k5_relat_1(X1,X0) = k3_xboole_0(X1,k2_zfmisc_1(X0,k10_xtuple_0(X1)))) [input t67_relat_1]
21. ! [X0] : (v1_relat_1(X0) => r1_tarski(X0,k2_zfmisc_1(k9_xtuple_0(X0),k10_xtuple_0(X0)))) [input t7_relat_1]
22. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,X1) => (r1_tarski(k2_zfmisc_1(X2,X0),k2_zfmisc_1(X2,X1)) & r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X2)))) [input t95_zfmisc_1]
23. ~! [X0,X1] : (v1_relat_1(X1) => (r1_tarski(k9_xtuple_0(X1),X0) => k5_relat_1(X1,X0) = X1)) [flattening 2]
24. ! [X0,X1] : (v1_relat_1(X1) => k5_relat_1(X1,X0) = k3_xboole_0(X1,k2_zfmisc_1(X0,k10_xtuple_0(X1)))) [flattening 20]
25. ! [X0,X1] : (r1_tarski(X0,X1) => k3_xboole_0(X0,X1) = X0) [flattening 18]
27. ! [X0,X1,X2] : (r1_tarski(X0,X1) => (r1_tarski(k2_zfmisc_1(X2,X0),k2_zfmisc_1(X2,X1)) & r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X2)))) [flattening 22]
28. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 17]
31. ? [X0,X1] : ((k5_relat_1(X1,X0) != X1 & r1_tarski(k9_xtuple_0(X1),X0)) & v1_relat_1(X1)) [ennf transformation 23]
32. ? [X0,X1] : (k5_relat_1(X1,X0) != X1 & r1_tarski(k9_xtuple_0(X1),X0) & v1_relat_1(X1)) [flattening 31]
33. ! [X0] : (r1_tarski(X0,k2_zfmisc_1(k9_xtuple_0(X0),k10_xtuple_0(X0))) | ~v1_relat_1(X0)) [ennf transformation 21]
35. ! [X0,X1] : (k5_relat_1(X1,X0) = k3_xboole_0(X1,k2_zfmisc_1(X0,k10_xtuple_0(X1))) | ~v1_relat_1(X1)) [ennf transformation 24]
36. ! [X0,X1] : (k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1)) [ennf transformation 25]
37. ! [X0,X1,X2] : ((r1_tarski(k2_zfmisc_1(X2,X0),k2_zfmisc_1(X2,X1)) & r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X2))) | ~r1_tarski(X0,X1)) [ennf transformation 27]
38. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 28]
39. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 38]
42. ? [X0,X1] : (k5_relat_1(X1,X0) != X1 & r1_tarski(k9_xtuple_0(X1),X0) & v1_relat_1(X1)) => (sK1 != k5_relat_1(sK1,sK0) & r1_tarski(k9_xtuple_0(sK1),sK0) & v1_relat_1(sK1)) [choice axiom]
43. sK1 != k5_relat_1(sK1,sK0) & r1_tarski(k9_xtuple_0(sK1),sK0) & v1_relat_1(sK1) [skolemisation 32,42]
47. v1_relat_1(sK1) [cnf transformation 43] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. r1_tarski(k9_xtuple_0(sK1),sK0) [cnf transformation 43] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. sK1 != k5_relat_1(sK1,sK0) [cnf transformation 43] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. r1_tarski(X0,k2_zfmisc_1(k9_xtuple_0(X0),k10_xtuple_0(X0))) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. k5_relat_1(X1,X0) = k3_xboole_0(X1,k2_zfmisc_1(X0,k10_xtuple_0(X1))) | ~v1_relat_1(X1) [cnf transformation 35] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
53. k3_xboole_0(X0,X1) = X0 | ~r1_tarski(X0,X1) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. r1_tarski(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X2)) | ~r1_tarski(X0,X1) [cnf transformation 37] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
58. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 39] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. 1 <=> sK1 = k5_relat_1(sK1,sK0) [avatar definition]
69. sK1 != k5_relat_1(sK1,sK0) <- (~1) [avatar component clause 67] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
70. ~1 [avatar split clause 49,67]
72. 2 <=> r1_tarski(k9_xtuple_0(sK1),sK0) [avatar definition]
74. r1_tarski(k9_xtuple_0(sK1),sK0) <- (2) [avatar component clause 72] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. 2 [avatar split clause 48,72]
77. 3 <=> v1_relat_1(sK1) [avatar definition]
79. v1_relat_1(sK1) <- (3) [avatar component clause 77] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. 3 [avatar split clause 47,77]
101. ~r1_tarski(X2,k2_zfmisc_1(X0,X3)) | r1_tarski(X2,k2_zfmisc_1(X1,X3)) | ~r1_tarski(X0,X1) [resolution 56,58] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
111. ~r1_tarski(X2,k2_zfmisc_1(X3,k10_xtuple_0(X2))) | ~v1_relat_1(X2) | k5_relat_1(X2,X3) = X2 [superposition 52,53] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
165. r1_tarski(X0,k2_zfmisc_1(X1,k10_xtuple_0(X0))) | ~r1_tarski(k9_xtuple_0(X0),X1) | ~v1_relat_1(X0) [resolution 101,50] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
281. ~r1_tarski(k9_xtuple_0(X0),X1) | ~v1_relat_1(X0) | ~v1_relat_1(X0) | k5_relat_1(X0,X1) = X0 [resolution 165,111] {a:3,w:13,wCS:5200,sine:1,thAx:0,allAx:5,thDist:-5}
287. k5_relat_1(X0,X1) = X0 | ~v1_relat_1(X0) | ~r1_tarski(k9_xtuple_0(X0),X1) [duplicate literal removal 281] {a:3,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
524. sK1 != sK1 | ~v1_relat_1(sK1) | ~r1_tarski(k9_xtuple_0(sK1),sK0) <- (~1) [superposition 69,287] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
527. ~v1_relat_1(sK1) | ~r1_tarski(k9_xtuple_0(sK1),sK0) <- (~1) [trivial inequality removal 524] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
529. ~r1_tarski(k9_xtuple_0(sK1),sK0) <- (~1, 3) [subsumption resolution 527,79] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
530. $false <- (~1, 2, 3) [subsumption resolution 529,74] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
531. 1 | ~2 | ~3 [avatar contradiction clause 530]
532. $false [avatar sat refutation 70,75,80,531] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:671,thDist:-671}
% SZS output end Proof for relat_1__t68_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 315
% Active clauses: 75
% Passive clauses: 173
% Final active clauses: 75
% Final passive clauses: 88
% Main loop iterations started: 85
% 
% >>> Simplifying Inferences
% Duplicate literals: 16
% Trivial inequalities: 1
% Fw subsumption resolutions: 12
% 
% >>> Deletion Inferences
% Simple tautologies: 16
% Equational tautologies: 1
% Forward subsumptions: 73
% 
% >>> Generating Inferences
% Binary resolution: 185
% Unit resulting resolution: 1
% Forward superposition: 13
% Backward superposition: 46
% 
% >>> AVATAR
% Split clauses: 20
% Split components: 40
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 35
% SAT solver unit clauses: 3
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 9722
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward demodulation: 0.001 s
% other: 0.004 s ( own 0.003 s ) 

