i: [1,96,0,2,1,0,0,2] rc1_ordinal6
a: [3,112,0,2,1,96]
s: 112
i: [1,97,0,2,1,0,0,2] rc1_ordinal1
a: [3,117,0,2,1,97]
s: 117
i: [1,99,0,2,1,0,0,255] rc2_ordinal2
a: [3,122,0,2,1,99]
s: 122
i: [1,98,0,2,1,0,0,255] rc2_ordinal2
a: [3,127,0,2,1,98]
s: 127
i: [1,101,0,2,1,0,0,2] rc1_funct_1
a: [3,132,0,2,1,101]
s: 132
i: [1,100,0,2,1,0,0,2] rc1_funct_1
a: [3,137,0,2,1,100]
s: 137
i: [1,68,0,9,1,1,0,0] t21_ordinal6
a: [3,107,0,9,1,68]
s: 107
i: [1,69,0,3,1,0,0,1] fc1_ordinal6
s: 69
i: [1,70,0,3,1,0,0,1] dt_k1_ordinal6
s: 70
i: [1,71,0,3,1,0,0,1] dt_k1_wellord2
s: 71
i: [1,74,0,3,1,0,0,255] dt_k2_ordinal6
s: 74
i: [1,75,0,3,1,0,0,255] dt_k2_ordinal6
s: 75
i: [1,76,0,3,1,0,0,255] fc1_wellord2
s: 76
i: [1,77,0,3,1,0,0,255] fc2_wellord2
s: 77
i: [1,78,0,3,1,0,0,255] fc3_wellord2
s: 78
i: [1,82,0,4,2,0,0,2] cc1_ordinal6
s: 82
d: [2,139,1,2,1,1,40,82,117]
a: [3,144,1,2,1,139]
s: 144
i: [1,83,0,4,2,0,0,255] cc1_ordinal1
s: 83
d: [2,146,1,2,1,1,40,83,117]
a: [3,151,1,2,1,146]
s: 151
i: [1,84,0,4,2,0,0,255] cc1_ordinal1
s: 84
i: [1,80,0,5,2,0,0,2] fc8_ordinal6
s: 80
d: [2,153,1,2,1,1,40,84,117]
a: [3,158,1,2,1,153]
s: 158
i: [1,81,0,5,2,0,0,2] fc8_ordinal6
s: 81
d: [2,140,1,3,1,0,40,82,70]
s: 140
i: [1,85,0,5,2,0,0,255] fc5_wellord2
s: 85
i: [1,86,0,5,2,0,0,255] fc4_wellord2
s: 86
d: [2,147,1,3,1,0,40,83,70]
s: 147
i: [1,87,0,5,2,0,0,2] dt_k2_wellord2
s: 87
d: [2,154,1,3,1,0,40,84,70]
s: 154
i: [1,93,0,6,3,0,0,255] cc2_ordinal1
s: 93
i: [1,72,0,7,1,0,0,1] d4_ordinal6
s: 72
d: [2,160,1,5,2,0,40,87,84]
s: 160
i: [1,79,0,8,2,0,0,2] d5_ordinal6
s: 79
d: [2,161,1,5,2,0,40,87,83]
s: 161
i: [1,94,0,8,3,0,0,2] dt_k3_wellord1
s: 94
i: [1,95,0,8,3,0,0,2] dt_k3_wellord1
s: 95
d: [2,162,1,5,2,0,40,87,82]
s: 162
i: [1,73,0,10,1,0,0,1] d6_ordinal6
s: 73
d: [2,168,1,9,2,0,44,79,72]
d: [2,172,1,6,1,0,21,168,69]
s: 172
i: [1,90,0,10,4,0,0,2] t40_wellord1
s: 90
i: [1,102,0,12,4,0,0,2] d2_wellord2
s: 102
d: [2,169,1,9,2,0,44,72,79]
d: [2,173,1,6,1,0,21,169,69]
s: 173
d: [2,185,0,9,3,0,21,102,87]
s: 185
d: [2,188,1,14,4,0,44,185,79]
d: [2,192,1,11,3,0,21,188,71]
d: [2,193,1,8,2,0,21,192,81]
s: 193
i: [1,89,0,14,5,0,0,2] d2_wellord2
s: 89
i: [1,92,0,22,8,0,0,1] d9_wellord1
s: 92
d: [2,195,2,10,2,0,44,193,172]
d: [2,198,2,7,1,0,21,195,69]
s: 198
i: [1,103,0,23,7,0,0,1] d9_wellord1
s: 103
d: [2,215,3,15,3,0,40,198,90]
d: [2,217,3,11,2,0,21,215,71]
d: [2,218,3,7,1,0,21,217,71]
s: 218
d: [2,219,0,19,6,0,21,103,95]
d: [2,220,0,15,5,0,21,219,94]
s: 220
d: [2,186,1,15,5,0,40,185,90]
d: [2,189,1,13,4,0,18,186]
d: [2,190,1,9,3,0,21,189,71]
s: 190
d: [2,194,2,15,4,0,40,193,90]
d: [2,196,2,11,3,0,21,194,71]
d: [2,197,2,8,2,0,21,196,71]
s: 197
d: [2,187,1,15,3,0,44,185,72]
d: [2,191,1,11,2,0,21,187,71]
s: 191
d: [2,234,2,14,4,0,44,190,79]
d: [2,238,2,11,3,0,21,234,71]
d: [2,239,2,8,2,0,21,238,81]
s: 239
d: [2,225,1,28,5,0,44,220,73]
d: [2,228,1,21,4,0,21,225,218]
d: [2,229,1,17,3,0,21,228,71]
d: [2,230,1,13,2,0,21,229,71]
s: 230
d: [2,202,1,35,8,0,44,92,73]
d: [2,210,1,31,7,0,21,202,71]
d: [2,211,1,27,6,0,21,210,71]
s: 211
d: [2,245,2,4,1,1,40,230,107]
a: [3,250,2,4,1,245]
s: 250
d: [2,253,3,0,0,1,52,140,250,81]
e: 253
d: [2,254,3,3,1,1,40,250,81]
d: [2,256,3,0,0,1,21,254,140]
e: 256
f: 256,245,68
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal6__t21_ordinal6
% SZS output start Proof for ordinal6__t21_ordinal6
1. ! [X0] : r3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0)),k2_ordinal6(X0)) [input t21_ordinal6]
2. ~! [X0] : r3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0)),k2_ordinal6(X0)) [negated conjecture 1]
4. ! [X0] : (v3_ordinal1(X0) => v1_ordinal6(X0)) [input cc1_ordinal6]
6. ! [X0] : (v1_relat_1(X0) => (v2_wellord1(X0) => ! [X1] : (v3_ordinal1(X1) => (k2_wellord2(X0) = X1 <=> r4_wellord1(X0,k1_wellord2(X1)))))) [input d2_wellord2]
7. ! [X0] : k1_ordinal6(X0) = k2_wellord2(k1_wellord2(k2_ordinal1(X0))) [input d4_ordinal6]
8. ! [X0] : (v1_ordinal6(X0) => k1_ordinal6(X0) = k2_wellord2(k1_wellord2(X0))) [input d5_ordinal6]
9. ! [X0] : k2_ordinal6(X0) = k3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0))) [input d6_ordinal6]
10. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ((r4_wellord1(X0,X1) & v2_wellord1(X0)) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (k3_wellord1(X0,X1) = X2 <=> r3_wellord1(X0,X1,X2)))))) [input d9_wellord1]
11. ! [X0] : v3_ordinal1(k1_ordinal6(X0)) [input dt_k1_ordinal6]
12. ! [X0] : v1_relat_1(k1_wellord2(X0)) [input dt_k1_wellord2]
15. ! [X0] : (v1_relat_1(X0) => v3_ordinal1(k2_wellord2(X0))) [input dt_k2_wellord2]
16. ! [X0,X1] : ((v1_relat_1(X1) & v1_relat_1(X0)) => (v1_funct_1(k3_wellord1(X0,X1)) & v1_relat_1(k3_wellord1(X0,X1)))) [input dt_k3_wellord1]
17. ! [X0] : v1_ordinal6(k2_ordinal1(X0)) [input fc1_ordinal6]
23. ! [X0] : (v1_ordinal6(X0) => (v2_wellord1(k1_wellord2(X0)) & v1_relat_1(k1_wellord2(X0)))) [input fc8_ordinal6]
28. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r4_wellord1(X0,X1) => r4_wellord1(X1,X0)))) [input t40_wellord1]
38. ? [X0] : ~r3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0)),k2_ordinal6(X0)) [ennf transformation 2]
39. ! [X0] : (k1_ordinal6(X0) = k2_wellord2(k1_wellord2(X0)) | ~v1_ordinal6(X0)) [ennf transformation 8]
40. ! [X0] : ((v2_wellord1(k1_wellord2(X0)) & v1_relat_1(k1_wellord2(X0))) | ~v1_ordinal6(X0)) [ennf transformation 23]
41. ! [X0] : (v1_ordinal6(X0) | ~v3_ordinal1(X0)) [ennf transformation 4]
45. ! [X0] : (v3_ordinal1(k2_wellord2(X0)) | ~v1_relat_1(X0)) [ennf transformation 15]
46. ! [X0] : ((! [X1] : ((k2_wellord2(X0) = X1 <=> r4_wellord1(X0,k1_wellord2(X1))) | ~v3_ordinal1(X1)) | ~v2_wellord1(X0)) | ~v1_relat_1(X0)) [ennf transformation 6]
47. ! [X0] : (! [X1] : ((k2_wellord2(X0) = X1 <=> r4_wellord1(X0,k1_wellord2(X1))) | ~v3_ordinal1(X1)) | ~v2_wellord1(X0) | ~v1_relat_1(X0)) [flattening 46]
48. ! [X0] : (! [X1] : ((r4_wellord1(X1,X0) | ~r4_wellord1(X0,X1)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 28]
49. ! [X0] : (! [X1] : (r4_wellord1(X1,X0) | ~r4_wellord1(X0,X1) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 48]
50. ! [X0] : (! [X1] : ((! [X2] : ((k3_wellord1(X0,X1) = X2 <=> r3_wellord1(X0,X1,X2)) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | (~r4_wellord1(X0,X1) | ~v2_wellord1(X0))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 10]
51. ! [X0] : (! [X1] : (! [X2] : ((k3_wellord1(X0,X1) = X2 <=> r3_wellord1(X0,X1,X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~r4_wellord1(X0,X1) | ~v2_wellord1(X0) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 50]
54. ! [X0,X1] : ((v1_funct_1(k3_wellord1(X0,X1)) & v1_relat_1(k3_wellord1(X0,X1))) | (~v1_relat_1(X1) | ~v1_relat_1(X0))) [ennf transformation 16]
55. ! [X0,X1] : ((v1_funct_1(k3_wellord1(X0,X1)) & v1_relat_1(k3_wellord1(X0,X1))) | ~v1_relat_1(X1) | ~v1_relat_1(X0)) [flattening 54]
56. ? [X0] : ~r3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0)),k2_ordinal6(X0)) => ~r3_wellord1(k1_wellord2(k1_ordinal6(sK0)),k1_wellord2(k2_ordinal1(sK0)),k2_ordinal6(sK0)) [choice axiom]
57. ~r3_wellord1(k1_wellord2(k1_ordinal6(sK0)),k1_wellord2(k2_ordinal1(sK0)),k2_ordinal6(sK0)) [skolemisation 38,56]
58. ! [X0] : (! [X1] : (((k2_wellord2(X0) = X1 | ~r4_wellord1(X0,k1_wellord2(X1))) & (r4_wellord1(X0,k1_wellord2(X1)) | k2_wellord2(X0) != X1)) | ~v3_ordinal1(X1)) | ~v2_wellord1(X0) | ~v1_relat_1(X0)) [nnf transformation 47]
59. ! [X0] : (! [X1] : (! [X2] : (((k3_wellord1(X0,X1) = X2 | ~r3_wellord1(X0,X1,X2)) & (r3_wellord1(X0,X1,X2) | k3_wellord1(X0,X1) != X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~r4_wellord1(X0,X1) | ~v2_wellord1(X0) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [nnf transformation 51]
68. ~r3_wellord1(k1_wellord2(k1_ordinal6(sK0)),k1_wellord2(k2_ordinal1(sK0)),k2_ordinal6(sK0)) [cnf transformation 57] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. v1_ordinal6(k2_ordinal1(X0)) [cnf transformation 17] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. v3_ordinal1(k1_ordinal6(X0)) [cnf transformation 11] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
71. v1_relat_1(k1_wellord2(X0)) [cnf transformation 12] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. k1_ordinal6(X0) = k2_wellord2(k1_wellord2(k2_ordinal1(X0))) [cnf transformation 7] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
73. k2_ordinal6(X0) = k3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0))) [cnf transformation 9] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. k1_ordinal6(X0) = k2_wellord2(k1_wellord2(X0)) | ~v1_ordinal6(X0) [cnf transformation 39] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
81. v2_wellord1(k1_wellord2(X0)) | ~v1_ordinal6(X0) [cnf transformation 40] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
82. ~v3_ordinal1(X0) | v1_ordinal6(X0) [cnf transformation 41] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
87. v3_ordinal1(k2_wellord2(X0)) | ~v1_relat_1(X0) [cnf transformation 45] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
88. r4_wellord1(X0,k1_wellord2(X1)) | k2_wellord2(X0) != X1 | ~v3_ordinal1(X1) | ~v2_wellord1(X0) | ~v1_relat_1(X0) [cnf transformation 58] {a:0,w:14,wCS:5600,sine:2,thAx:0,allAx:1,thDist:-1}
90. ~r4_wellord1(X0,X1) | r4_wellord1(X1,X0) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 49] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
91. r3_wellord1(X0,X1,X2) | k3_wellord1(X0,X1) != X2 | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~r4_wellord1(X0,X1) | ~v2_wellord1(X0) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 59] {a:0,w:22,wCS:8800,sine:1,thAx:0,allAx:1,thDist:-1}
94. v1_relat_1(k3_wellord1(X0,X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 55] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
95. v1_funct_1(k3_wellord1(X0,X1)) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 55] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
102. r4_wellord1(X0,k1_wellord2(k2_wellord2(X0))) | ~v3_ordinal1(k2_wellord2(X0)) | ~v2_wellord1(X0) | ~v1_relat_1(X0) [equality resolution 88] {a:0,w:12,wCS:4800,sine:2,thAx:0,allAx:1,thDist:-1}
103. r3_wellord1(X0,X1,k3_wellord1(X0,X1)) | ~v1_funct_1(k3_wellord1(X0,X1)) | ~v1_relat_1(k3_wellord1(X0,X1)) | ~r4_wellord1(X0,X1) | ~v2_wellord1(X0) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [equality resolution 91] {a:0,w:23,wCS:9200,sine:1,thAx:0,allAx:1,thDist:-1}
105. 1 <=> r3_wellord1(k1_wellord2(k1_ordinal6(sK0)),k1_wellord2(k2_ordinal1(sK0)),k2_ordinal6(sK0)) [avatar definition]
107. ~r3_wellord1(k1_wellord2(k1_ordinal6(sK0)),k1_wellord2(k2_ordinal1(sK0)),k2_ordinal6(sK0)) <- (~1) [avatar component clause 105] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
108. ~1 [avatar split clause 68,105]
140. v1_ordinal6(k1_ordinal6(X0)) [resolution 82,70] {a:1,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
168. k1_ordinal6(X0) = k1_ordinal6(k2_ordinal1(X0)) | ~v1_ordinal6(k2_ordinal1(X0)) [superposition 79,72] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
172. k1_ordinal6(X0) = k1_ordinal6(k2_ordinal1(X0)) [subsumption resolution 168,69] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
185. r4_wellord1(X0,k1_wellord2(k2_wellord2(X0))) | ~v2_wellord1(X0) | ~v1_relat_1(X0) [subsumption resolution 102,87] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
188. r4_wellord1(k1_wellord2(X1),k1_wellord2(k1_ordinal6(X1))) | ~v2_wellord1(k1_wellord2(X1)) | ~v1_relat_1(k1_wellord2(X1)) | ~v1_ordinal6(X1) [superposition 185,79] {a:1,w:14,wCS:5600,sine:2,thAx:0,allAx:2,thDist:-2}
192. r4_wellord1(k1_wellord2(X1),k1_wellord2(k1_ordinal6(X1))) | ~v2_wellord1(k1_wellord2(X1)) | ~v1_ordinal6(X1) [subsumption resolution 188,71] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
193. r4_wellord1(k1_wellord2(X1),k1_wellord2(k1_ordinal6(X1))) | ~v1_ordinal6(X1) [subsumption resolution 192,81] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
195. r4_wellord1(k1_wellord2(k2_ordinal1(X0)),k1_wellord2(k1_ordinal6(X0))) | ~v1_ordinal6(k2_ordinal1(X0)) [superposition 193,172] {a:2,w:10,wCS:4000,sine:1,thAx:0,allAx:4,thDist:-4}
198. r4_wellord1(k1_wellord2(k2_ordinal1(X0)),k1_wellord2(k1_ordinal6(X0))) [subsumption resolution 195,69] {a:2,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
215. r4_wellord1(k1_wellord2(k1_ordinal6(X1)),k1_wellord2(k2_ordinal1(X1))) | ~v1_relat_1(k1_wellord2(k1_ordinal6(X1))) | ~v1_relat_1(k1_wellord2(k2_ordinal1(X1))) [resolution 198,90] {a:3,w:15,wCS:6000,sine:1,thAx:0,allAx:5,thDist:-5}
217. r4_wellord1(k1_wellord2(k1_ordinal6(X1)),k1_wellord2(k2_ordinal1(X1))) | ~v1_relat_1(k1_wellord2(k2_ordinal1(X1))) [subsumption resolution 215,71] {a:3,w:11,wCS:4400,sine:1,thAx:0,allAx:5,thDist:-5}
218. r4_wellord1(k1_wellord2(k1_ordinal6(X1)),k1_wellord2(k2_ordinal1(X1))) [subsumption resolution 217,71] {a:3,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
219. r3_wellord1(X0,X1,k3_wellord1(X0,X1)) | ~v1_relat_1(k3_wellord1(X0,X1)) | ~r4_wellord1(X0,X1) | ~v2_wellord1(X0) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [subsumption resolution 103,95] {a:0,w:19,wCS:7600,sine:1,thAx:0,allAx:1,thDist:-1}
220. r3_wellord1(X0,X1,k3_wellord1(X0,X1)) | ~r4_wellord1(X0,X1) | ~v2_wellord1(X0) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [subsumption resolution 219,94] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
225. r3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0)),k2_ordinal6(X0)) | ~r4_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0))) | ~v2_wellord1(k1_wellord2(k1_ordinal6(X0))) | ~v1_relat_1(k1_wellord2(k2_ordinal1(X0))) | ~v1_relat_1(k1_wellord2(k1_ordinal6(X0))) [superposition 220,73] {a:1,w:28,wCS:11200,sine:1,thAx:0,allAx:2,thDist:-2}
228. r3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0)),k2_ordinal6(X0)) | ~v2_wellord1(k1_wellord2(k1_ordinal6(X0))) | ~v1_relat_1(k1_wellord2(k2_ordinal1(X0))) | ~v1_relat_1(k1_wellord2(k1_ordinal6(X0))) [subsumption resolution 225,218] {a:1,w:21,wCS:8400,sine:1,thAx:0,allAx:2,thDist:-2}
229. r3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0)),k2_ordinal6(X0)) | ~v2_wellord1(k1_wellord2(k1_ordinal6(X0))) | ~v1_relat_1(k1_wellord2(k1_ordinal6(X0))) [subsumption resolution 228,71] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
230. r3_wellord1(k1_wellord2(k1_ordinal6(X0)),k1_wellord2(k2_ordinal1(X0)),k2_ordinal6(X0)) | ~v2_wellord1(k1_wellord2(k1_ordinal6(X0))) [subsumption resolution 229,71] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
245. ~v2_wellord1(k1_wellord2(k1_ordinal6(sK0))) <- (~1) [resolution 230,107] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
248. 11 <=> v2_wellord1(k1_wellord2(k1_ordinal6(sK0))) [avatar definition]
250. ~v2_wellord1(k1_wellord2(k1_ordinal6(sK0))) <- (~11) [avatar component clause 248] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
251. ~11 | 1 [avatar split clause 245,105,248]
254. ~v1_ordinal6(k1_ordinal6(sK0)) <- (~11) [resolution 250,81] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
256. $false <- (~11) [subsumption resolution 254,140] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
257. 11 [avatar contradiction clause 256]
258. $false [avatar sat refutation 108,251,257] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:54,thDist:-54}
% SZS output end Proof for ordinal6__t21_ordinal6
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 27
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 7
% 
% >>> Saturation
% Initial clauses: 34
% Generated clauses: 142
% Active clauses: 45
% Passive clauses: 67
% Final active clauses: 45
% Final passive clauses: 10
% Main loop iterations started: 57
% 
% >>> Simplifying Inferences
% Duplicate literals: 16
% Fw subsumption resolutions: 30
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 33
% 
% >>> Generating Inferences
% Binary resolution: 26
% Unit resulting resolution: 1
% Forward superposition: 18
% Backward superposition: 14
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 9
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% splitting component index maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

