i: [1,224,0,2,1,1,0,0] t137_xreal_1
a: [3,353,0,2,1,224]
s: 353
i: [1,223,0,2,1,1,0,0] t137_xreal_1
a: [3,358,0,2,1,223]
s: 358
i: [1,226,0,3,1,1,0,0] t137_xreal_1
a: [3,343,0,3,1,226]
s: 343
i: [1,225,0,3,1,1,0,0] t137_xreal_1
a: [3,348,0,3,1,225]
s: 348
i: [1,227,0,5,1,1,0,0] t137_xreal_1
a: [3,338,0,5,1,227]
s: 338
i: [1,228,0,2,1,0,0,255] fc1_numbers
a: [3,363,0,2,1,228]
s: 363
i: [1,229,0,2,1,0,0,2] fc1_xboole_0
a: [3,368,0,2,1,229]
s: 368
i: [1,236,0,2,1,0,0,3] fc6_ordinal1
a: [3,398,0,2,1,236]
s: 398
i: [1,235,0,2,1,0,0,3] fc6_ordinal1
a: [3,403,0,2,1,235]
s: 403
i: [1,302,0,2,1,0,0,2] rc2_xboole_0
a: [3,415,0,2,1,302]
s: 415
i: [1,303,0,2,1,0,0,4] rc4_ordinal1
a: [3,420,0,2,1,303]
s: 420
i: [1,304,0,2,1,0,0,4] rc1_ordinal1
a: [3,425,0,2,1,304]
s: 425
i: [1,305,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,430,0,2,1,305]
s: 430
i: [1,306,0,2,1,0,0,1] rc1_xreal_0
a: [3,435,0,2,1,306]
s: 435
i: [1,307,0,2,1,0,0,2] rc1_xxreal_0
a: [3,440,0,2,1,307]
s: 440
i: [1,308,0,2,1,0,0,2] rc1_xboole_0
a: [3,445,0,2,1,308]
s: 445
i: [1,312,0,2,1,0,0,255] rc2_ordinal1
a: [3,450,0,2,1,312]
s: 450
i: [1,311,0,2,1,0,0,255] rc2_ordinal1
a: [3,455,0,2,1,311]
s: 455
i: [1,310,0,2,1,0,0,255] rc2_ordinal1
a: [3,460,0,2,1,310]
s: 460
i: [1,309,0,2,1,0,0,255] rc2_ordinal1
a: [3,465,0,2,1,309]
s: 465
i: [1,314,0,2,1,0,0,4] rc5_ordinal1
a: [3,470,0,2,1,314]
s: 470
i: [1,313,0,2,1,0,0,4] rc5_ordinal1
a: [3,475,0,2,1,313]
s: 475
i: [1,316,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,480,0,2,1,316]
s: 480
i: [1,315,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,485,0,2,1,315]
s: 485
i: [1,320,0,2,1,0,0,2] rc2_xreal_0
a: [3,490,0,2,1,320]
s: 490
i: [1,319,0,2,1,0,0,2] rc2_xreal_0
a: [3,495,0,2,1,319]
s: 495
i: [1,318,0,2,1,0,0,2] rc2_xreal_0
a: [3,500,0,2,1,318]
s: 500
i: [1,317,0,2,1,0,0,2] rc2_xreal_0
a: [3,505,0,2,1,317]
s: 505
i: [1,324,0,2,1,0,0,2] rc3_xreal_0
a: [3,510,0,2,1,324]
s: 510
i: [1,323,0,2,1,0,0,2] rc3_xreal_0
a: [3,515,0,2,1,323]
s: 515
i: [1,322,0,2,1,0,0,2] rc3_xreal_0
a: [3,520,0,2,1,322]
s: 520
i: [1,321,0,2,1,0,0,2] rc3_xreal_0
a: [3,525,0,2,1,321]
s: 525
i: [1,326,0,2,1,0,0,2] rc2_xxreal_0
a: [3,530,0,2,1,326]
s: 530
i: [1,325,0,2,1,0,0,2] rc2_xxreal_0
a: [3,535,0,2,1,325]
s: 535
i: [1,328,0,2,1,0,0,2] rc3_xxreal_0
a: [3,540,0,2,1,328]
s: 540
i: [1,327,0,2,1,0,0,2] rc3_xxreal_0
a: [3,545,0,2,1,327]
s: 545
i: [1,332,0,2,1,0,0,2] rc4_xreal_0
a: [3,550,0,2,1,332]
s: 550
i: [1,331,0,2,1,0,0,2] rc4_xreal_0
a: [3,555,0,2,1,331]
s: 555
i: [1,330,0,2,1,0,0,2] rc4_xreal_0
a: [3,560,0,2,1,330]
s: 560
i: [1,329,0,2,1,0,0,2] rc4_xreal_0
a: [3,565,0,2,1,329]
s: 565
i: [1,334,0,2,1,0,0,2] rc4_xxreal_0
a: [3,570,0,2,1,334]
s: 570
i: [1,333,0,2,1,0,0,2] rc4_xxreal_0
a: [3,575,0,2,1,333]
s: 575
i: [1,287,0,3,1,0,0,255] reflexivity_r1_tarski
s: 287
i: [1,230,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,373,0,3,1,230]
s: 373
i: [1,231,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,378,0,3,1,231]
s: 378
r: 368
d: [2,577,0,2,1,2,23,368,378]
a: [3,581,0,2,1,577]
s: 581
i: [1,232,0,3,1,0,0,2] t1_numerals
a: [3,383,0,3,1,232]
s: 383
d: [2,583,0,3,1,2,22,383,378]
a: [3,587,0,3,1,583]
s: 587
i: [1,238,0,4,2,0,0,4] cc6_ordinal1
s: 238
i: [1,239,0,4,2,0,0,4] cc2_xcmplx_0
s: 239
i: [1,240,0,4,2,0,0,4] cc2_xreal_0
s: 240
d: [2,589,1,2,1,1,40,238,420]
a: [3,594,1,2,1,589]
s: 594
i: [1,241,0,4,2,0,0,4] cc2_xxreal_0
s: 241
d: [2,590,1,2,1,1,40,238,470]
a: [3,599,1,2,1,590]
s: 599
i: [1,242,0,4,2,0,0,255] cc1_ordinal1
s: 242
i: [1,243,0,4,2,0,0,255] cc1_ordinal1
s: 243
d: [2,601,1,2,1,1,40,239,420]
a: [3,606,1,2,1,601]
s: 606
i: [1,246,0,4,2,0,0,2] cc3_xreal_0
s: 246
d: [2,687,1,2,1,1,40,246,358]
a: [3,696,1,2,1,687]
s: 696
i: [1,247,0,4,2,0,0,2] cc4_xreal_0
s: 247
i: [1,256,0,4,2,0,0,255] cc9_ordinal1
s: 256
d: [2,688,1,2,1,1,40,246,353]
a: [3,701,1,2,1,688]
s: 701
i: [1,257,0,4,2,0,0,4] cc7_ordinal1
s: 257
d: [2,708,1,2,1,1,40,247,358]
a: [3,717,1,2,1,708]
s: 717
i: [1,258,0,4,2,0,0,4] cc3_ordinal1
s: 258
i: [1,286,0,4,1,0,0,3] existence_m1_subset_1
s: 286
d: [2,709,1,2,1,1,40,247,353]
a: [3,722,1,2,1,709]
s: 722
i: [1,233,0,4,1,0,0,255] dt_k5_numbers
a: [3,388,0,4,1,233]
s: 388
d: [2,602,1,2,1,1,40,239,470]
a: [3,611,1,2,1,602]
s: 611
i: [1,234,0,4,1,0,0,255] dt_k6_numbers
a: [3,393,0,4,1,234]
s: 393
d: [2,801,0,4,1,2,22,388,373]
a: [3,805,0,4,1,801]
s: 805
d: [2,613,1,2,1,1,40,240,420]
a: [3,618,1,2,1,613]
s: 618
d: [2,807,0,4,1,2,22,393,373]
a: [3,811,0,4,1,807]
s: 811
d: [2,614,1,2,1,1,40,240,470]
a: [3,623,1,2,1,614]
s: 623
i: [1,259,0,5,2,0,0,2] t6_boole
s: 259
i: [1,260,0,5,2,0,0,3] cc8_ordinal1
s: 260
d: [2,625,1,2,1,1,40,241,420]
a: [3,630,1,2,1,625]
s: 630
i: [1,261,0,5,2,0,0,255] cc1_xcmplx_0
s: 261
d: [2,626,1,2,1,1,40,241,470]
a: [3,635,1,2,1,626]
s: 635
i: [1,262,0,5,2,0,0,255] cc1_xreal_0
s: 262
i: [1,283,0,5,2,0,0,255] t7_boole
s: 283
d: [2,637,1,2,1,1,40,242,398]
a: [3,645,1,2,1,637]
s: 645
i: [1,300,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,410,0,5,2,300]
s: 410
d: [2,639,1,2,1,1,40,242,425]
a: [3,655,1,2,1,639]
s: 655
d: [2,819,0,5,2,1,22,259,378]
s: 819
i: [1,267,0,6,3,0,0,255] cc2_ordinal1
s: 267
d: [2,842,1,5,2,2,44,343,819]
s: 842
i: [1,270,0,6,3,0,0,2] cc5_xxreal_0
s: 270
d: [2,843,1,5,2,2,44,348,819]
s: 843
i: [1,272,0,6,3,0,0,2] cc5_xxreal_0
s: 272
i: [1,273,0,6,3,0,0,2] cc3_xxreal_0
s: 273
d: [2,841,1,7,2,2,44,338,819]
s: 841
i: [1,275,0,6,3,0,0,2] cc3_xxreal_0
s: 275
d: [2,890,2,8,2,3,40,841,410]
a: [3,899,2,4,1,890]
s: 899
i: [1,277,0,6,3,0,0,2] cc7_xxreal_0
s: 277
i: [1,278,0,6,3,0,0,2] cc7_xxreal_0
s: 278
d: [2,662,1,2,1,1,40,243,398]
a: [3,670,1,2,1,662]
s: 670
i: [1,279,0,6,2,0,0,255] t1_subset
s: 279
d: [2,664,1,2,1,1,40,243,425]
a: [3,680,1,2,1,664]
s: 680
i: [1,288,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 288
i: [1,244,0,7,3,0,0,4] cc5_ordinal1
s: 244
d: [2,689,1,2,1,1,40,246,435]
a: [3,706,1,2,1,689]
s: 706
i: [1,245,0,7,2,0,0,1] t5_arithm
s: 245
d: [2,710,1,2,1,1,40,247,435]
a: [3,727,1,2,1,710]
s: 727
i: [1,281,0,7,2,0,0,255] t3_subset
s: 281
i: [1,282,0,7,3,0,0,2] t8_boole
s: 282
d: [2,729,1,2,1,1,40,256,581]
a: [3,736,1,2,1,729]
s: 736
i: [1,237,0,8,3,0,0,255] cc10_ordinal1
s: 237
d: [2,730,1,2,1,1,40,256,445]
a: [3,741,1,2,1,730]
s: 741
i: [1,264,0,8,4,0,0,2] cc4_xxreal_0
s: 264
i: [1,266,0,8,4,0,0,2] cc6_xxreal_0
s: 266
d: [2,891,2,6,2,3,40,841,843]
a: [3,903,2,2,1,891]
s: 903
r: 903
i: [1,268,0,8,4,0,0,2] cc8_xxreal_0
s: 268
a: [3,904,2,2,1,891]
s: 904
i: [1,280,0,8,3,0,0,255] t2_subset
s: 280
i: [1,298,0,8,3,0,0,1] fc5_xcmplx_0
s: 298
d: [2,928,2,6,3,1,40,264,717]
a: [3,943,2,2,1,928]
s: 943
i: [1,299,0,8,3,0,0,1] fc8_xreal_0
s: 299
d: [2,929,2,6,3,1,40,264,722]
a: [3,952,2,2,1,929]
s: 952
r: 952
i: [1,285,0,9,3,0,0,255] t5_subset
s: 285
i: [1,284,0,10,3,0,0,255] t4_subset
s: 284
a: [3,951,2,2,1,929]
s: 951
i: [1,301,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 301
a: [3,956,2,2,1,929]
s: 956
i: [1,250,0,11,5,0,0,1] t2_real
s: 250
i: [1,251,0,11,5,0,0,1] t1_real
s: 251
d: [2,1096,1,8,4,1,40,250,343]
d: [2,1103,1,6,3,2,21,1096,353]
a: [3,1107,1,2,1,1103]
s: 1107
i: [1,252,0,11,5,0,0,1] t8_real
s: 252
a: [3,946,2,2,1,928]
s: 946
i: [1,253,0,11,5,0,0,1] t7_real
s: 253
i: [1,254,0,11,5,0,0,1] t4_real
s: 254
a: [3,960,2,2,1,929]
s: 960
d: [2,1161,3,0,0,3,52,722,956,960,272]
e: 1161
d: [2,1162,3,4,2,1,40,960,272]
d: [2,1165,3,2,1,2,21,1162,956]
d: [2,1166,3,0,0,3,21,1165,722]
e: 1166
r: 960
i: [1,255,0,11,5,0,0,1] t3_real
s: 255
a: [3,895,2,4,1,890]
s: 895
i: [1,289,0,12,5,0,0,1] fc31_xreal_0
s: 289
i: [1,290,0,12,5,0,0,1] fc32_xreal_0
s: 290
d: [2,1145,2,4,2,2,44,1107,819]
s: 1145
d: [2,1170,3,0,0,4,52,550,565,1145]
e: 1170
d: [2,1173,3,2,1,3,40,1145,565]
d: [2,1176,3,0,0,4,21,1173,550]
e: 1176
r: 956
r: 1145
r: 1107
i: [1,291,0,12,5,0,0,1] fc34_xreal_0
s: 291
a: [3,1106,1,2,1,1103]
s: 1106
i: [1,292,0,12,5,0,0,1] fc33_xreal_0
s: 292
i: [1,297,0,12,5,0,0,1] fc9_xcmplx_0
s: 297
d: [2,1078,1,11,3,1,40,301,338]
a: [3,1084,1,2,1,1078]
s: 1084
i: [1,248,0,13,6,0,0,1] t6_real
s: 248
a: [3,1111,1,2,1,1103]
s: 1111
i: [1,249,0,13,6,0,0,1] t5_real
s: 249
i: [1,296,0,14,4,0,0,255] existence_m2_subset_1
s: 296
d: [2,1127,1,8,4,1,40,251,343]
d: [2,1134,1,6,3,2,21,1127,956]
d: [2,1135,1,4,2,3,21,1134,353]
a: [3,1138,1,2,1,1135]
s: 1138
i: [1,293,0,15,5,0,0,255] dt_m2_subset_1
s: 293
a: [3,955,2,2,1,929]
s: 955
i: [1,294,0,15,5,0,0,255] redefinition_m2_subset_1
s: 294
i: [1,295,0,15,5,0,0,255] redefinition_m2_subset_1
s: 295
s: 960
d: [2,731,1,2,1,1,40,256,565]
a: [3,746,1,2,1,731]
s: 746
d: [2,1195,2,2,1,1,40,1106,246]
a: [3,1201,2,2,1,1195]
s: 1201
d: [2,732,1,2,1,1,40,256,575]
a: [3,751,1,2,1,732]
s: 751
d: [2,753,1,2,1,1,40,257,581]
a: [3,760,1,2,1,753]
s: 760
d: [2,1147,2,14,5,2,40,252,841]
a: [3,1155,2,4,1,1147]
s: 1155
d: [2,1271,3,0,0,3,52,358,353,1155,299]
e: 1271
d: [2,1272,3,4,2,1,40,1155,299]
d: [2,1274,3,2,1,2,21,1272,353]
d: [2,1275,3,0,0,3,21,1274,358]
e: 1275
r: 1111
r: 1138
r: 1155
a: [3,1110,1,2,1,1103]
s: 1110
d: [2,1277,2,0,0,3,52,1084,581,1110,270]
e: 1277
d: [2,1279,2,4,2,1,40,1110,270]
d: [2,1283,2,2,1,2,21,1279,581]
d: [2,1284,2,0,0,3,21,1283,1084]
e: 1284
r: 1110
s: 1111
s: 1138
d: [2,754,1,2,1,1,40,257,445]
a: [3,765,1,2,1,754]
s: 765
a: [3,1154,2,4,1,1147]
s: 1154
d: [2,1308,3,0,0,2,52,895,1154,247]
e: 1308
d: [2,1309,3,4,1,1,40,1154,247]
d: [2,1312,3,0,0,2,21,1309,895]
e: 1312
r: 895
a: [3,1089,1,5,1,1078]
s: 1089
a: [3,894,2,4,1,890]
s: 894
d: [2,755,1,2,1,1,40,257,565]
a: [3,770,1,2,1,755]
s: 770
d: [2,756,1,2,1,1,40,257,575]
a: [3,775,1,2,1,756]
s: 775
a: [3,1159,2,4,1,1147]
s: 1159
d: [2,1361,3,0,0,5,52,946,358,955,353,1159,289]
e: 1361
d: [2,1363,3,8,4,1,40,1159,289]
d: [2,1366,3,6,3,2,21,1363,353]
d: [2,1367,3,4,2,3,21,1366,955]
d: [2,1368,3,2,1,4,21,1367,358]
d: [2,1369,3,0,0,5,21,1368,946]
e: 1369
r: 955
r: 960
r: 1138
a: [3,1139,1,2,1,1135]
s: 1139
d: [2,1374,2,0,0,3,52,1084,581,1139,273]
e: 1374
d: [2,1375,2,4,2,1,40,1139,273]
d: [2,1378,2,2,1,2,21,1375,581]
d: [2,1379,2,0,0,3,21,1378,1084]
e: 1379
d: [2,1098,1,8,4,1,40,250,348]
d: [2,1118,1,6,3,2,21,1098,358]
d: [2,1183,1,6,3,2,21,1127,353]
d: [2,1146,1,12,4,1,40,252,338]
d: [2,1194,2,2,1,1,40,1106,247]
d: [2,1371,3,4,2,3,21,1366,358]
d: [2,1372,3,2,1,4,21,1371,946]
f: 1379,1372,1284,1275,1194,1146,1183,1176,1118,577,329,332,231,229,223,224,225,226,227
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t137_xreal_1
% SZS output start Proof for xreal_1__t137_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0)) => r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers)))) [input t137_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0)) => r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers)))) [negated conjecture 1]
12. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
15. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
19. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v2_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc31_xreal_0]
30. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
32. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
40. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
58. ? [X0] : (v1_xreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0) & v1_xboole_0(X0)) [input rc4_xreal_0]
73. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
79. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
82. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
86. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
87. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
100. ? [X0] : (? [X1] : ((~r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & (r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
101. ? [X0] : (? [X1] : (~r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 100]
111. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 86]
116. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 32]
117. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 116]
118. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 30]
119. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 118]
120. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 40]
121. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 120]
131. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 79]
143. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 15]
144. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 143]
145. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 12]
146. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 145]
159. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 19]
160. ! [X0,X1] : (~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 159]
177. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 87]
178. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 177]
183. ? [X0] : (? [X1] : (~r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~r1_xxreal_0(k7_xcmplx_0(sK0,X1),k6_numbers) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
184. ? [X1] : (~r1_xxreal_0(k7_xcmplx_0(sK0,X1),k6_numbers) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) => (~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) & r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) [choice axiom]
185. (~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) & r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 101,184,183]
219. ? [X0] : (v1_xreal_0(X0) & v1_xxreal_0(X0) & v1_xcmplx_0(X0) & v1_xboole_0(X0)) => (v1_xreal_0(sK18) & v1_xxreal_0(sK18) & v1_xcmplx_0(sK18) & v1_xboole_0(sK18)) [choice axiom]
220. v1_xreal_0(sK18) & v1_xxreal_0(sK18) & v1_xcmplx_0(sK18) & v1_xboole_0(sK18) [skolemisation 58,219]
223. v1_xreal_0(sK0) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
224. v1_xreal_0(sK1) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
225. r1_xxreal_0(k6_numbers,sK0) [cnf transformation 185] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
226. r1_xxreal_0(sK1,k6_numbers) [cnf transformation 185] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
227. ~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) [cnf transformation 185] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
229. v1_xboole_0(k1_xboole_0) [cnf transformation 73] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
231. k6_numbers = k1_xboole_0 [cnf transformation 82] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
247. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 111] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
250. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 117] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
251. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 119] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
252. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 121] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
259. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 131] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
270. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 144] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
273. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 146] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
289. ~v2_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 160] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
299. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 178] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
329. v1_xboole_0(sK18) [cnf transformation 220] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
332. v1_xreal_0(sK18) [cnf transformation 220] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
336. 1 <=> r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) [avatar definition]
338. ~r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) <- (~1) [avatar component clause 336] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
339. ~1 [avatar split clause 227,336]
341. 2 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
343. r1_xxreal_0(sK1,k6_numbers) <- (2) [avatar component clause 341] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
344. 2 [avatar split clause 226,341]
346. 3 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
348. r1_xxreal_0(k6_numbers,sK0) <- (3) [avatar component clause 346] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
349. 3 [avatar split clause 225,346]
351. 4 <=> v1_xreal_0(sK1) [avatar definition]
353. v1_xreal_0(sK1) <- (4) [avatar component clause 351] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
354. 4 [avatar split clause 224,351]
356. 5 <=> v1_xreal_0(sK0) [avatar definition]
358. v1_xreal_0(sK0) <- (5) [avatar component clause 356] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
359. 5 [avatar split clause 223,356]
366. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
368. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 366] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
369. 7 [avatar split clause 229,366]
376. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
378. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 376] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
379. 9 [avatar split clause 231,376]
548. 44 <=> v1_xreal_0(sK18) [avatar definition]
550. v1_xreal_0(sK18) <- (44) [avatar component clause 548] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
551. 44 [avatar split clause 332,548]
563. 47 <=> v1_xboole_0(sK18) [avatar definition]
565. v1_xboole_0(sK18) <- (47) [avatar component clause 563] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
566. 47 [avatar split clause 329,563]
577. v1_xboole_0(k6_numbers) <- (7, 9) [backward demodulation 368,378] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
579. 50 <=> v1_xboole_0(k6_numbers) [avatar definition]
581. v1_xboole_0(k6_numbers) <- (50) [avatar component clause 579] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
582. 50 | ~7 | ~9 [avatar split clause 577,376,366,579]
819. k6_numbers = X0 | ~v1_xboole_0(X0) <- (9) [forward demodulation 259,378] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
945. 103 <=> v3_xxreal_0(sK0) [avatar definition]
946. ~v3_xxreal_0(sK0) <- (~103) [avatar component clause 945] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
954. 105 <=> v2_xxreal_0(sK1) [avatar definition]
1083. 127 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1084. v1_xxreal_0(k6_numbers) <- (127) [avatar component clause 1083] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1098. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (3) [resolution 250,348] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1105. 130 <=> v1_xreal_0(k6_numbers) [avatar definition]
1106. v1_xreal_0(k6_numbers) <- (130) [avatar component clause 1105] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1107. ~v1_xreal_0(k6_numbers) <- (~130) [avatar component clause 1105] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1109. 131 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1110. v3_xxreal_0(k6_numbers) <- (131) [avatar component clause 1109] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1118. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 5) [subsumption resolution 1098,358] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1119. ~130 | 131 | ~103 | ~3 | ~5 [avatar split clause 1118,356,346,945,1109,1105]
1127. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (2) [resolution 251,343] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1137. 134 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1139. v2_xxreal_0(k6_numbers) <- (134) [avatar component clause 1137] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1145. ~v1_xboole_0(X0) | ~v1_xreal_0(X0) <- (9, ~130) [superposition 1107,819] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1146. v3_xxreal_0(k6_numbers) | v2_xxreal_0(k7_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~1) [resolution 252,338] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1153. 135 <=> v1_xreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1155. ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~135) [avatar component clause 1153] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1157. 136 <=> v2_xxreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1159. v2_xxreal_0(k7_xcmplx_0(sK0,sK1)) <- (136) [avatar component clause 1157] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1173. ~v1_xreal_0(sK18) <- (9, 47, ~130) [resolution 1145,565] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
1176. $false <- (9, 44, 47, ~130) [subsumption resolution 1173,550] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
1177. ~9 | ~44 | ~47 | 130 [avatar contradiction clause 1176]
1183. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 4) [subsumption resolution 1127,353] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1184. ~130 | 134 | ~105 | ~2 | ~4 [avatar split clause 1183,351,341,954,1137,1105]
1190. ~135 | ~130 | 136 | 131 | 1 [avatar split clause 1146,336,1109,1157,1105,1153]
1194. v1_xxreal_0(k6_numbers) <- (130) [resolution 1106,247] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1197. 127 | ~130 [avatar split clause 1194,1105,1083]
1272. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~135) [resolution 1155,299] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1274. ~v1_xreal_0(sK0) <- (4, ~135) [subsumption resolution 1272,353] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1275. $false <- (4, 5, ~135) [subsumption resolution 1274,358] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1276. ~4 | ~5 | 135 [avatar contradiction clause 1275]
1279. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (131) [resolution 1110,270] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1283. ~v1_xxreal_0(k6_numbers) <- (50, 131) [subsumption resolution 1279,581] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1284. $false <- (50, 127, 131) [subsumption resolution 1283,1084] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1285. ~50 | ~127 | ~131 [avatar contradiction clause 1284]
1363. ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (136) [resolution 1159,289] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1366. v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (4, 136) [subsumption resolution 1363,353] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1371. v2_xxreal_0(sK1) | v3_xxreal_0(sK0) <- (4, 5, 136) [subsumption resolution 1366,358] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1372. v2_xxreal_0(sK1) <- (4, 5, ~103, 136) [subsumption resolution 1371,946] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1373. 105 | ~4 | ~5 | 103 | ~136 [avatar split clause 1372,1157,945,356,351,954]
1375. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (134) [resolution 1139,273] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1378. ~v1_xxreal_0(k6_numbers) <- (50, 134) [subsumption resolution 1375,581] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1379. $false <- (50, 127, 134) [subsumption resolution 1378,1084] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1380. ~50 | ~127 | ~134 [avatar contradiction clause 1379]
1381. $false [avatar sat refutation 339,344,349,354,359,369,379,551,566,582,1119,1177,1184,1190,1197,1276,1285,1373,1380] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:947,thDist:-947}
% SZS output end Proof for xreal_1__t137_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 87
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 112
% Generated clauses: 560
% Active clauses: 156
% Passive clauses: 242
% Final active clauses: 141
% Final passive clauses: 69
% Main loop iterations started: 168
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 60
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Forward subsumptions: 64
% 
% >>> Generating Inferences
% Binary resolution: 180
% Unit resulting resolution: 7
% Factoring: 1
% Forward superposition: 21
% Backward superposition: 6
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 29
% Split components: 74
% Sat splitting refutations: 14
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 167
% SAT solver unit clauses: 50
% SAT solver binary clauses: 61
% 
% Memory used [KB]: 5756
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% unit clause index maintenance: 0.001 s
% forward subsumption resolution: 0.001 s
% backward subsumption: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

