i: [1,21,0,3,1,1,0,0] t17_xtuple_0
a: [3,36,0,3,1,21]
s: 36
i: [1,22,0,5,1,1,0,0] t17_xtuple_0
a: [3,31,0,5,1,22]
s: 31
i: [1,27,0,3,1,0,0,1] reflexivity_r1_tarski
s: 27
i: [1,23,0,6,1,0,0,1] d17_xtuple_0
s: 23
i: [1,24,0,6,1,0,0,255] d14_xtuple_0
s: 24
i: [1,25,0,8,2,0,0,2] t9_xtuple_0
s: 25
d: [2,38,1,7,1,0,44,24,24]
s: 38
i: [1,26,0,8,2,0,0,2] t10_xtuple_0
s: 26
d: [2,41,2,8,1,0,44,24,38]
s: 41
d: [2,39,1,9,2,0,44,25,23]
s: 39
d: [2,40,1,9,2,0,44,25,23]
s: 40
d: [2,42,3,9,1,0,44,41,24]
s: 42
d: [2,48,2,10,2,0,44,39,23]
s: 48
d: [2,54,3,8,2,0,40,48,26]
s: 54
d: [2,56,4,0,0,2,52,36,31,54]
e: 56
d: [2,57,4,3,1,1,40,54,31]
d: [2,59,4,0,0,2,21,57,36]
e: 59
f: 59,21,22
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t17_xtuple_0
% SZS output start Proof for xtuple_0__t17_xtuple_0
1. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1))) [input t17_xtuple_0]
2. ~! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1))) [negated conjecture 1]
4. ! [X0] : k14_xtuple_0(X0) = k10_xtuple_0(k11_xtuple_0(X0)) [input d17_xtuple_0]
10. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1))) [input t10_xtuple_0]
11. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1))) [input t9_xtuple_0]
12. ~! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1))) [flattening 2]
13. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1))) [flattening 11]
14. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1))) [flattening 10]
16. ? [X0,X1] : (~r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1)) & r1_tarski(X0,X1)) [ennf transformation 12]
17. ! [X0,X1] : (r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 13]
18. ! [X0,X1] : (r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 14]
19. ? [X0,X1] : (~r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1)) & r1_tarski(X0,X1)) => (~r1_tarski(k14_xtuple_0(sK0),k14_xtuple_0(sK1)) & r1_tarski(sK0,sK1)) [choice axiom]
20. ~r1_tarski(k14_xtuple_0(sK0),k14_xtuple_0(sK1)) & r1_tarski(sK0,sK1) [skolemisation 16,19]
21. r1_tarski(sK0,sK1) [cnf transformation 20] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
22. ~r1_tarski(k14_xtuple_0(sK0),k14_xtuple_0(sK1)) [cnf transformation 20] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
23. k14_xtuple_0(X0) = k10_xtuple_0(k11_xtuple_0(X0)) [cnf transformation 4] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
25. r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 17] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
26. r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 18] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
29. 1 <=> r1_tarski(k14_xtuple_0(sK0),k14_xtuple_0(sK1)) [avatar definition]
31. ~r1_tarski(k14_xtuple_0(sK0),k14_xtuple_0(sK1)) <- (~1) [avatar component clause 29] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
32. ~1 [avatar split clause 22,29]
34. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
36. r1_tarski(sK0,sK1) <- (2) [avatar component clause 34] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
37. 2 [avatar split clause 21,34]
39. r1_tarski(k14_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(k11_xtuple_0(X0),X1) [superposition 25,23] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
48. ~r1_tarski(k11_xtuple_0(X1),k11_xtuple_0(X0)) | r1_tarski(k14_xtuple_0(X1),k14_xtuple_0(X0)) [superposition 39,23] {a:2,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
54. r1_tarski(k14_xtuple_0(X0),k14_xtuple_0(X1)) | ~r1_tarski(X0,X1) [resolution 48,26] {a:3,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
57. ~r1_tarski(sK0,sK1) <- (~1) [resolution 54,31] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
59. $false <- (~1, 2) [subsumption resolution 57,36] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
60. 1 | ~2 [avatar contradiction clause 59]
61. $false [avatar sat refutation 32,37,60] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:20,thDist:-20}
% SZS output end Proof for xtuple_0__t17_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 7
% Generated clauses: 30
% Active clauses: 14
% Passive clauses: 20
% Final active clauses: 14
% Final passive clauses: 6
% Main loop iterations started: 14
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 1
% 
% >>> Deletion Inferences
% Forward subsumptions: 1
% Fw demodulations to eq. taut.: 4
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 1
% Forward superposition: 8
% Backward superposition: 6
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 6
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

