i: [1,67,0,3,1,1,0,0] t4_subset
a: [3,99,0,3,1,67]
s: 99
i: [1,65,0,3,1,1,0,0] t4_subset
a: [3,109,0,3,1,65]
s: 109
i: [1,66,0,4,1,1,0,0] t4_subset
a: [3,104,0,4,1,66]
s: 104
i: [1,68,0,3,1,0,0,1] fc1_subset_1
s: 68
i: [1,91,0,3,1,0,0,1] rc2_subset_1
s: 91
i: [1,94,0,3,1,0,0,255] reflexivity_r1_tarski
s: 94
i: [1,89,0,4,1,0,0,1] existence_m1_subset_1
s: 89
i: [1,93,0,4,1,0,0,255] rc3_subset_1
s: 93
i: [1,72,0,5,2,0,0,1] rc1_subset_1
s: 72
i: [1,75,0,5,2,0,0,255] t6_boole
s: 75
i: [1,88,0,5,2,0,0,1] t7_boole
s: 88
d: [2,112,1,2,1,1,40,88,109]
a: [3,116,1,2,1,112]
s: 116
i: [1,90,0,5,1,0,0,1] rc2_subset_1
s: 90
i: [1,92,0,5,1,0,0,255] rc3_subset_1
s: 92
i: [1,74,0,6,2,0,0,255] rc4_subset_1
s: 74
i: [1,95,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 95
d: [2,118,1,3,1,1,40,95,109]
a: [3,122,1,3,1,118]
s: 122
i: [1,71,0,7,2,0,0,1] rc1_subset_1
s: 71
i: [1,73,0,7,2,0,0,255] rc4_subset_1
s: 73
i: [1,80,0,7,3,0,0,1] d1_subset_1
s: 80
i: [1,81,0,7,3,0,0,1] d1_subset_1
s: 81
d: [2,130,1,4,2,1,40,81,99]
a: [3,136,1,2,1,130]
s: 136
i: [1,85,0,7,2,0,0,255] t3_subset
s: 85
d: [2,143,1,3,1,1,40,85,104]
a: [3,152,1,3,1,143]
s: 152
i: [1,86,0,7,2,0,0,255] t3_subset
s: 86
i: [1,87,0,7,3,0,0,255] t8_boole
s: 87
d: [2,147,1,4,1,0,40,85,90]
s: 147
i: [1,76,0,8,3,0,0,1] cc1_subset_1
s: 76
d: [2,148,1,4,1,0,40,85,92]
s: 148
i: [1,78,0,8,3,0,0,1] d1_subset_1
s: 78
i: [1,79,0,8,3,0,0,1] d1_subset_1
s: 79
d: [2,166,1,7,2,1,40,78,104]
d: [2,173,1,4,1,1,21,166,68]
a: [3,177,1,4,1,173]
s: 177
d: [2,183,0,6,2,0,21,79,88]
s: 183
d: [2,191,1,3,1,1,40,183,109]
a: [3,196,1,3,1,191]
s: 196
i: [1,83,0,8,2,0,0,255] d3_tarski
s: 83
i: [1,84,0,8,2,0,0,255] d3_tarski
s: 84
d: [2,184,2,4,1,1,40,177,95]
a: [3,189,2,4,1,184]
s: 189
i: [1,77,0,9,3,0,0,255] cc4_subset_1
s: 77
d: [2,127,1,5,2,0,40,80,89]
s: 127
i: [1,82,0,9,3,0,0,255] d3_tarski
s: 82
i: [1,69,0,11,4,0,0,255] cc3_subset_1
s: 69
d: [2,214,2,6,2,1,40,82,152]
s: 214
i: [1,70,0,11,4,0,0,255] cc2_subset_1
s: 70
d: [2,230,3,3,1,2,40,214,109]
a: [3,235,3,3,1,230]
s: 235
d: [2,237,4,0,0,2,52,99,235,183]
e: 237
d: [2,238,4,3,1,1,40,235,183]
d: [2,242,4,0,0,2,21,238,99]
e: 242
f: 242,230,143,65,66,67
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset__t4_subset
% SZS output start Proof for subset__t4_subset
1. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
2. ~! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [negated conjecture 1]
8. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
9. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
20. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
22. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
24. ~! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 2]
25. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 8]
26. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 9]
27. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 20]
29. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 22]
31. ? [X0,X1,X2] : (~m1_subset_1(X0,X2) & (m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1))) [ennf transformation 24]
32. ? [X0,X1,X2] : (~m1_subset_1(X0,X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) [flattening 31]
42. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 25]
43. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 26]
45. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 29]
47. ? [X0,X1,X2] : (~m1_subset_1(X0,X2) & m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => (~m1_subset_1(sK0,sK2) & m1_subset_1(sK1,k1_zfmisc_1(sK2)) & r2_hidden(sK0,sK1)) [choice axiom]
48. ~m1_subset_1(sK0,sK2) & m1_subset_1(sK1,k1_zfmisc_1(sK2)) & r2_hidden(sK0,sK1) [skolemisation 32,47]
53. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 42]
54. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 43]
55. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 54]
56. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK5(X0,X1),X1) & r2_hidden(sK5(X0,X1),X0))) [choice axiom]
57. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK5(X0,X1),X1) & r2_hidden(sK5(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 55,56]
58. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 27]
65. r2_hidden(sK0,sK1) [cnf transformation 48] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
66. m1_subset_1(sK1,k1_zfmisc_1(sK2)) [cnf transformation 48] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
67. ~m1_subset_1(sK0,sK2) [cnf transformation 48] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
79. m1_subset_1(X1,X0) | ~r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 53] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
82. ~r1_tarski(X0,X1) | ~r2_hidden(X3,X0) | r2_hidden(X3,X1) [cnf transformation 57] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
85. ~m1_subset_1(X0,k1_zfmisc_1(X1)) | r1_tarski(X0,X1) [cnf transformation 58] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
88. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 45] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
97. 1 <=> m1_subset_1(sK0,sK2) [avatar definition]
99. ~m1_subset_1(sK0,sK2) <- (~1) [avatar component clause 97] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
100. ~1 [avatar split clause 67,97]
102. 2 <=> m1_subset_1(sK1,k1_zfmisc_1(sK2)) [avatar definition]
104. m1_subset_1(sK1,k1_zfmisc_1(sK2)) <- (2) [avatar component clause 102] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
105. 2 [avatar split clause 66,102]
107. 3 <=> r2_hidden(sK0,sK1) [avatar definition]
109. r2_hidden(sK0,sK1) <- (3) [avatar component clause 107] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
110. 3 [avatar split clause 65,107]
143. r1_tarski(sK1,sK2) <- (2) [resolution 85,104] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
150. 8 <=> r1_tarski(sK1,sK2) [avatar definition]
152. r1_tarski(sK1,sK2) <- (8) [avatar component clause 150] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
153. 8 | ~2 [avatar split clause 143,102,150]
183. ~r2_hidden(X1,X0) | m1_subset_1(X1,X0) [subsumption resolution 79,88] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
214. ~r2_hidden(X2,sK1) | r2_hidden(X2,sK2) <- (8) [resolution 82,152] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
230. r2_hidden(sK0,sK2) <- (3, 8) [resolution 214,109] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
233. 12 <=> r2_hidden(sK0,sK2) [avatar definition]
235. r2_hidden(sK0,sK2) <- (12) [avatar component clause 233] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
236. 12 | ~3 | ~8 [avatar split clause 230,150,107,233]
238. m1_subset_1(sK0,sK2) <- (12) [resolution 235,183] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
242. $false <- (~1, 12) [subsumption resolution 238,99] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
243. 1 | ~12 [avatar contradiction clause 242]
244. $false [avatar sat refutation 100,105,110,153,236,243] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:74,thDist:-74}
% SZS output end Proof for subset__t4_subset
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 31
% Generated clauses: 129
% Active clauses: 42
% Passive clauses: 70
% Final active clauses: 42
% Final passive clauses: 26
% Main loop iterations started: 44
% 
% >>> Simplifying Inferences
% Duplicate literals: 5
% Fw subsumption resolutions: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 70
% Unit resulting resolution: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 3
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

