i: [1,39,0,8,2,1,0,0] t33_zfmisc_1
a: [3,57,0,4,1,39]
s: 57
a: [3,60,0,4,1,39]
s: 60
r: 57
d: [2,94,0,3,1,2,23,57,60]
a: [3,98,0,3,1,94]
s: 98
i: [1,40,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,72,0,2,1,40]
s: 72
i: [1,41,0,2,1,0,0,1] fc1_xboole_0
a: [3,77,0,2,1,41]
s: 77
i: [1,50,0,2,1,0,0,2] rc2_xboole_0
a: [3,87,0,2,1,50]
s: 87
i: [1,51,0,2,1,0,0,2] rc1_xboole_0
a: [3,92,0,2,1,51]
s: 92
i: [1,43,0,3,1,0,0,1] fc2_xboole_0
s: 43
d: [2,100,1,2,1,1,44,43,60]
a: [3,104,1,2,1,100]
s: 104
i: [1,49,0,3,1,0,0,1] reflexivity_r1_tarski
s: 49
d: [2,107,1,0,0,1,40,49,98]
e: 107
d: [2,106,1,0,0,1,52,98,49]
e: 106
r: 60
r: 98
i: [1,38,0,7,2,1,0,0] t33_zfmisc_1
a: [3,66,0,3,1,38]
s: 66
a: [3,56,0,4,1,39]
s: 56
a: [3,61,0,4,1,39]
s: 61
i: [1,42,0,3,1,0,0,255] d2_xboole_0
a: [3,82,0,3,1,42]
s: 82
r: 72
i: [1,53,0,4,1,0,0,1] l3_zfmisc_1
s: 53
i: [1,44,0,5,2,0,0,1] t6_boole
s: 44
i: [1,52,0,5,1,0,0,1] l3_zfmisc_1
s: 52
i: [1,48,0,7,3,0,0,1] t8_boole
s: 48
d: [2,113,1,6,2,0,44,53,44]
s: 113
i: [1,45,0,11,3,0,0,1] l3_zfmisc_1
s: 45
d: [2,115,1,0,0,3,52,66,61,56,45]
e: 115
d: [2,118,1,7,2,1,40,45,56]
d: [2,121,1,4,1,2,21,118,66]
d: [2,122,1,0,0,3,21,121,61]
e: 122
r: 56
r: 66
r: 104
a: [3,65,0,3,1,38]
s: 65
r: 61
s: 57
d: [2,130,0,4,1,2,22,57,65]
d: [2,131,0,0,0,2,21,130,53]
e: 131
i: [1,37,0,11,3,1,0,0] t33_zfmisc_1
f: 131,122,106,94,37,38,39
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t33_zfmisc_1
% SZS output start Proof for zfmisc_1__t33_zfmisc_1
1. ! [X0] : ! [X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [input t33_zfmisc_1]
2. ~! [X0] : ! [X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [negated conjecture 1]
9. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
16. ! [X0] : ! [X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [input l3_zfmisc_1]
17. ~! [X0,X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [flattening 2]
18. ! [X0,X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [flattening 16]
21. ! [X0] : r1_tarski(X0,X0) [rectify 9]
24. ? [X0,X1] : (r1_tarski(X0,k1_tarski(X1)) <~> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [ennf transformation 17]
27. ? [X0,X1] : (((k1_tarski(X1) != X0 & k1_xboole_0 != X0) | ~r1_tarski(X0,k1_tarski(X1))) & ((k1_tarski(X1) = X0 | k1_xboole_0 = X0) | r1_tarski(X0,k1_tarski(X1)))) [nnf transformation 24]
28. ? [X0,X1] : (((k1_tarski(X1) != X0 & k1_xboole_0 != X0) | ~r1_tarski(X0,k1_tarski(X1))) & (k1_tarski(X1) = X0 | k1_xboole_0 = X0 | r1_tarski(X0,k1_tarski(X1)))) [flattening 27]
29. ? [X0,X1] : (((k1_tarski(X1) != X0 & k1_xboole_0 != X0) | ~r1_tarski(X0,k1_tarski(X1))) & (k1_tarski(X1) = X0 | k1_xboole_0 = X0 | r1_tarski(X0,k1_tarski(X1)))) => (((sK0 != k1_tarski(sK1) & k1_xboole_0 != sK0) | ~r1_tarski(sK0,k1_tarski(sK1))) & (sK0 = k1_tarski(sK1) | k1_xboole_0 = sK0 | r1_tarski(sK0,k1_tarski(sK1)))) [choice axiom]
30. ((sK0 != k1_tarski(sK1) & k1_xboole_0 != sK0) | ~r1_tarski(sK0,k1_tarski(sK1))) & (sK0 = k1_tarski(sK1) | k1_xboole_0 = sK0 | r1_tarski(sK0,k1_tarski(sK1))) [skolemisation 28,29]
31. ! [X0,X1] : ((r1_tarski(X0,k1_tarski(X1)) | (k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & ((k1_tarski(X1) = X0 | k1_xboole_0 = X0) | ~r1_tarski(X0,k1_tarski(X1)))) [nnf transformation 18]
32. ! [X0,X1] : ((r1_tarski(X0,k1_tarski(X1)) | (k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & (k1_tarski(X1) = X0 | k1_xboole_0 = X0 | ~r1_tarski(X0,k1_tarski(X1)))) [flattening 31]
37. sK0 = k1_tarski(sK1) | k1_xboole_0 = sK0 | r1_tarski(sK0,k1_tarski(sK1)) [cnf transformation 30] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. k1_xboole_0 != sK0 | ~r1_tarski(sK0,k1_tarski(sK1)) [cnf transformation 30] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. sK0 != k1_tarski(sK1) | ~r1_tarski(sK0,k1_tarski(sK1)) [cnf transformation 30] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
45. ~r1_tarski(X0,k1_tarski(X1)) | k1_xboole_0 = X0 | k1_tarski(X1) = X0 [cnf transformation 32] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
46. r1_tarski(X0,k1_tarski(X1)) | k1_xboole_0 != X0 [cnf transformation 32] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
49. r1_tarski(X0,X0) [cnf transformation 21] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. r1_tarski(k1_xboole_0,k1_tarski(X1)) [equality resolution 46] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. 1 <=> r1_tarski(sK0,k1_tarski(sK1)) [avatar definition]
56. r1_tarski(sK0,k1_tarski(sK1)) <- (1) [avatar component clause 55] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
57. ~r1_tarski(sK0,k1_tarski(sK1)) <- (~1) [avatar component clause 55] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
59. 2 <=> sK0 = k1_tarski(sK1) [avatar definition]
60. sK0 = k1_tarski(sK1) <- (2) [avatar component clause 59] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
61. sK0 != k1_tarski(sK1) <- (~2) [avatar component clause 59] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. ~1 | ~2 [avatar split clause 39,59,55]
64. 3 <=> k1_xboole_0 = sK0 [avatar definition]
65. k1_xboole_0 = sK0 <- (3) [avatar component clause 64] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66. k1_xboole_0 != sK0 <- (~3) [avatar component clause 64] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
67. ~1 | ~3 [avatar split clause 38,64,55]
68. 1 | 3 | 2 [avatar split clause 37,59,64,55]
94. ~r1_tarski(sK0,sK0) <- (~1, 2) [backward demodulation 57,60] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
96. 9 <=> r1_tarski(sK0,sK0) [avatar definition]
98. ~r1_tarski(sK0,sK0) <- (~9) [avatar component clause 96] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
99. ~9 | 1 | ~2 [avatar split clause 94,59,55,96]
106. $false <- (~9) [unit resulting resolution 98,49] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
109. 9 [avatar contradiction clause 106]
118. k1_xboole_0 = sK0 | sK0 = k1_tarski(sK1) <- (1) [resolution 45,56] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
121. sK0 = k1_tarski(sK1) <- (1, ~3) [subsumption resolution 118,66] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
122. $false <- (1, ~2, ~3) [subsumption resolution 121,61] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
123. ~1 | 2 | 3 [avatar contradiction clause 122]
130. ~r1_tarski(k1_xboole_0,k1_tarski(sK1)) <- (~1, 3) [forward demodulation 57,65] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
131. $false <- (~1, 3) [subsumption resolution 130,53] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
132. 1 | ~3 [avatar contradiction clause 131]
133. $false [avatar sat refutation 62,67,68,99,109,123,132] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:59,thDist:-59}
% SZS output end Proof for zfmisc_1__t33_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 49
% Active clauses: 20
% Passive clauses: 24
% Final active clauses: 12
% Final passive clauses: 2
% Main loop iterations started: 22
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 3
% Fw demodulations: 1
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Equational tautologies: 2
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 5
% Unit resulting resolution: 2
% Forward superposition: 1
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 19
% SAT solver unit clauses: 7
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

