i: [1,107,0,2,1,1,0,0] t8_ordinal3
a: [3,167,0,2,1,107]
s: 167
i: [1,109,0,3,1,1,0,0] t8_ordinal3
a: [3,157,0,3,1,109]
s: 157
i: [1,108,0,3,1,1,0,0] t8_ordinal3
a: [3,162,0,3,1,108]
s: 162
i: [1,110,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,172,0,2,1,110]
s: 172
i: [1,111,0,2,1,0,0,1] fc1_xboole_0
a: [3,177,0,2,1,111]
s: 177
i: [1,144,0,2,1,0,0,2] rc2_xboole_0
a: [3,194,0,2,1,144]
s: 194
i: [1,145,0,2,1,0,0,255] rc4_ordinal1
a: [3,199,0,2,1,145]
s: 199
i: [1,146,0,2,1,0,0,2] rc1_xboole_0
a: [3,204,0,2,1,146]
s: 204
i: [1,147,0,2,1,0,0,1] rc1_ordinal1
a: [3,209,0,2,1,147]
s: 209
i: [1,151,0,2,1,0,0,255] rc2_ordinal1
a: [3,214,0,2,1,151]
s: 214
i: [1,150,0,2,1,0,0,255] rc2_ordinal1
a: [3,219,0,2,1,150]
s: 219
i: [1,149,0,2,1,0,0,255] rc2_ordinal1
a: [3,224,0,2,1,149]
s: 224
i: [1,148,0,2,1,0,0,255] rc2_ordinal1
a: [3,229,0,2,1,148]
s: 229
i: [1,153,0,2,1,0,0,255] rc5_ordinal1
a: [3,234,0,2,1,153]
s: 234
i: [1,152,0,2,1,0,0,255] rc5_ordinal1
a: [3,239,0,2,1,152]
s: 239
i: [1,113,0,3,1,0,0,1] t2_xboole_1
s: 113
i: [1,136,0,3,1,0,0,255] irreflexivity_r2_xboole_0
s: 136
i: [1,137,0,3,1,0,0,2] reflexivity_r1_tarski
s: 137
i: [1,112,0,3,1,0,0,255] d2_xboole_0
a: [3,182,0,3,1,112]
s: 182
r: 172
i: [1,116,0,4,2,0,0,255] cc6_ordinal1
s: 116
i: [1,117,0,4,2,0,0,255] cc9_ordinal1
s: 117
d: [2,242,1,2,1,1,40,116,199]
a: [3,247,1,2,1,242]
s: 247
i: [1,118,0,4,2,0,0,255] cc7_ordinal1
s: 118
d: [2,243,1,2,1,1,40,116,234]
a: [3,252,1,2,1,243]
s: 252
i: [1,119,0,4,2,0,0,1] cc3_ordinal1
s: 119
i: [1,121,0,4,2,0,0,255] cc1_ordinal1
s: 121
d: [2,290,1,2,1,1,40,121,167]
a: [3,298,1,2,1,290]
s: 298
i: [1,122,0,4,2,0,0,255] cc1_ordinal1
s: 122
d: [2,315,1,2,1,1,40,122,167]
a: [3,323,1,2,1,315]
s: 323
i: [1,135,0,4,1,0,0,2] existence_m1_subset_1
s: 135
i: [1,120,0,5,2,0,0,1] t6_boole
s: 120
d: [2,343,1,5,2,1,44,157,120]
s: 343
i: [1,132,0,5,2,0,0,1] t7_boole
s: 132
d: [2,254,1,2,1,1,40,117,204]
a: [3,259,1,2,1,254]
s: 259
i: [1,140,0,7,3,0,0,255] reflexivity_r1_ordinal1
a: [3,189,0,5,2,140]
s: 189
i: [1,125,0,6,3,0,0,255] cc2_ordinal1
s: 125
d: [2,255,1,2,1,1,40,117,177]
a: [3,264,1,2,1,255]
s: 264
i: [1,126,0,6,2,0,0,1] t1_subset
s: 126
d: [2,266,1,2,1,1,40,118,204]
a: [3,271,1,2,1,266]
s: 271
i: [1,138,0,6,2,0,0,255] antisymmetry_r2_xboole_0
s: 138
i: [1,139,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 139
d: [2,267,1,2,1,1,40,118,177]
a: [3,276,1,2,1,267]
s: 276
i: [1,123,0,7,3,0,0,2] cc5_ordinal1
s: 123
d: [2,278,1,2,1,1,40,119,204]
a: [3,283,1,2,1,278]
s: 283
i: [1,124,0,7,3,0,0,2] cc2_ordinal2
s: 124
i: [1,129,0,7,2,0,0,255] t3_subset
s: 129
d: [2,279,1,2,1,1,40,119,177]
a: [3,288,1,2,1,279]
s: 288
i: [1,130,0,7,2,0,0,255] t3_subset
s: 130
d: [2,292,1,2,1,1,40,121,209]
a: [3,308,1,2,1,292]
s: 308
i: [1,131,0,7,3,0,0,1] t8_boole
s: 131
i: [1,114,0,8,3,0,0,255] cc10_ordinal1
s: 114
d: [2,317,1,2,1,1,40,122,209]
a: [3,333,1,2,1,317]
s: 333
i: [1,127,0,8,3,0,0,1] t2_subset
s: 127
d: [2,291,2,2,1,1,40,121,247]
a: [3,303,2,2,1,291]
s: 303
i: [1,128,0,9,3,0,0,255] d8_xboole_0
s: 128
i: [1,134,0,9,3,0,0,255] t5_subset
s: 134
d: [2,294,2,2,1,1,40,121,252]
a: [3,313,2,2,1,294]
s: 313
i: [1,115,0,10,4,0,0,255] t11_ordinal1
s: 115
d: [2,316,2,2,1,1,40,122,247]
a: [3,328,2,2,1,316]
s: 328
i: [1,133,0,10,3,0,0,255] t4_subset
s: 133
i: [1,141,0,10,4,0,0,255] connectedness_r1_ordinal1
s: 141
d: [2,319,2,2,1,1,40,122,252]
a: [3,338,2,2,1,319]
s: 338
i: [1,142,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 142
d: [2,353,2,2,1,1,40,283,122]
a: [3,358,2,2,1,353]
s: 358
i: [1,143,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 143
d: [2,341,1,5,2,0,44,113,120]
s: 341
d: [2,354,2,2,1,1,40,283,121]
a: [3,363,2,2,1,354]
s: 363
d: [2,352,1,5,2,0,40,123,135]
s: 352
d: [2,366,2,2,1,1,40,288,122]
a: [3,372,2,2,1,366]
s: 372
d: [2,365,1,5,1,0,40,129,135]
s: 365
d: [2,401,1,7,3,0,40,143,113]
d: [2,404,1,5,2,1,21,401,288]
s: 404
d: [2,367,2,2,1,1,40,288,121]
a: [3,377,2,2,1,367]
s: 377
d: [2,381,1,6,2,0,40,114,135]
s: 381
d: [2,411,3,4,2,1,44,372,120]
s: 411
d: [2,384,1,6,2,0,40,127,135]
s: 384
d: [2,386,1,6,2,0,40,128,113]
s: 386
d: [2,420,2,5,2,1,40,384,343]
a: [3,427,2,3,1,420]
s: 427
d: [2,340,1,7,3,0,44,120,120]
s: 340
d: [2,417,3,4,2,1,44,377,120]
s: 417
d: [2,382,1,7,3,0,40,114,130]
s: 382
d: [2,388,1,7,2,0,40,134,135]
s: 388
d: [2,408,2,5,2,0,40,352,122]
s: 408
d: [2,380,1,8,3,0,40,130,123]
s: 380
d: [2,409,2,5,2,0,40,352,121]
s: 409
d: [2,389,1,8,3,0,40,134,130]
s: 389
d: [2,391,1,8,2,0,40,133,135]
s: 391
d: [2,450,2,5,2,0,40,389,113]
a: [3,456,2,2,1,450]
s: 456
r: 239
r: 229
r: 194
r: 427
r: 417
r: 411
r: 389
r: 388
r: 343
r: 341
r: 134
r: 132
r: 131
r: 120
r: 119
r: 118
r: 117
d: [2,464,3,0,0,2,40,456,177]
e: 464
d: [2,463,3,0,0,2,40,456,204]
e: 463
d: [2,462,3,0,0,2,52,204,456]
e: 462
r: 456
s: 194
s: 427
s: 229
s: 239
s: 343
s: 117
a: [3,459,2,3,1,450]
s: 459
s: 118
s: 119
s: 411
s: 120
s: 417
s: 132
s: 131
d: [2,486,3,4,2,1,44,377,120]
s: 486
s: 134
s: 341
d: [2,479,1,5,2,0,44,113,120]
s: 479
s: 388
d: [2,447,2,5,2,0,40,380,137]
s: 447
d: [2,478,1,7,3,0,44,120,120]
s: 478
d: [2,489,3,5,2,1,44,459,120]
s: 489
d: [2,493,1,7,2,0,40,134,135]
s: 493
s: 389
d: [2,421,2,6,2,0,40,384,139]
s: 421
d: [2,494,1,8,3,0,40,134,130]
s: 494
d: [2,435,2,6,2,0,40,386,138]
s: 435
d: [2,392,1,9,3,0,40,133,130]
s: 392
d: [2,385,1,10,3,0,40,127,130]
s: 385
d: [2,444,2,6,2,0,40,388,384]
s: 444
d: [2,397,1,14,6,0,40,142,141]
d: [2,398,1,10,4,0,18,397]
s: 398
d: [2,500,2,6,2,0,40,388,384]
s: 500
d: [2,405,2,9,4,0,40,341,143]
d: [2,407,2,7,3,0,21,405,119]
s: 407
d: [2,415,2,7,3,1,44,404,120]
s: 415
d: [2,517,3,6,2,0,40,444,119]
s: 517
d: [2,488,2,7,3,1,44,404,120]
s: 488
d: [2,518,3,6,2,0,40,444,118]
s: 518
d: [2,497,2,9,4,0,40,341,143]
d: [2,499,2,7,3,0,21,497,119]
s: 499
d: [2,406,2,8,3,0,40,341,128]
s: 406
d: [2,519,3,6,2,0,40,444,117]
s: 519
d: [2,436,2,8,3,0,44,386,120]
s: 436
d: [2,515,4,6,2,1,40,444,417]
s: 515
d: [2,434,2,10,4,0,40,386,115]
d: [2,437,2,8,3,1,21,434,377]
s: 437
d: [2,539,3,0,0,4,52,167,162,157,437]
e: 539
d: [2,544,3,5,2,2,40,437,157]
d: [2,551,3,2,1,3,21,544,162]
d: [2,552,3,0,0,4,21,551,167]
e: 552
d: [2,545,3,7,3,2,40,437,343]
d: [2,554,3,4,2,3,21,545,162]
d: [2,555,3,2,1,4,21,554,167]
d: [2,556,3,0,0,5,21,555,177]
e: 556
f: 556,367,279,111,107,108,109
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t8_ordinal3
% SZS output start Proof for ordinal3__t8_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => (k1_xboole_0 != X0 => r2_hidden(k1_xboole_0,X0))) [input t8_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => (k1_xboole_0 != X0 => r2_hidden(k1_xboole_0,X0))) [negated conjecture 1]
6. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
9. ! [X0] : (v1_xboole_0(X0) => v3_ordinal1(X0)) [input cc3_ordinal1]
17. ! [X0] : ! [X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [input d8_xboole_0]
23. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
34. ! [X0] : (v1_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_xboole_0(X0,X1) => r2_hidden(X0,X1)))) [input t11_ordinal1]
37. ! [X0] : r1_tarski(k1_xboole_0,X0) [input t2_xboole_1]
41. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
46. ! [X0,X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [flattening 17]
54. ! [X0,X1] : ((X0 != X1 & r1_tarski(X0,X1)) => r2_xboole_0(X0,X1)) [unused predicate definition removal 46]
56. ? [X0] : ((~r2_hidden(k1_xboole_0,X0) & k1_xboole_0 != X0) & v3_ordinal1(X0)) [ennf transformation 2]
57. ? [X0] : (~r2_hidden(k1_xboole_0,X0) & k1_xboole_0 != X0 & v3_ordinal1(X0)) [flattening 56]
59. ! [X0] : (! [X1] : ((r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1)) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [ennf transformation 34]
60. ! [X0] : (! [X1] : (r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [flattening 59]
64. ! [X0] : (v3_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 9]
65. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 41]
66. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 6]
74. ! [X0,X1] : (r2_xboole_0(X0,X1) | (X0 = X1 | ~r1_tarski(X0,X1))) [ennf transformation 54]
75. ! [X0,X1] : (r2_xboole_0(X0,X1) | X0 = X1 | ~r1_tarski(X0,X1)) [flattening 74]
89. ? [X0] : (~r2_hidden(k1_xboole_0,X0) & k1_xboole_0 != X0 & v3_ordinal1(X0)) => (~r2_hidden(k1_xboole_0,sK0) & k1_xboole_0 != sK0 & v3_ordinal1(sK0)) [choice axiom]
90. ~r2_hidden(k1_xboole_0,sK0) & k1_xboole_0 != sK0 & v3_ordinal1(sK0) [skolemisation 57,89]
107. v3_ordinal1(sK0) [cnf transformation 90] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
108. k1_xboole_0 != sK0 [cnf transformation 90] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
109. ~r2_hidden(k1_xboole_0,sK0) [cnf transformation 90] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
111. v1_xboole_0(k1_xboole_0) [cnf transformation 23] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
113. r1_tarski(k1_xboole_0,X0) [cnf transformation 37] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
115. ~r2_xboole_0(X0,X1) | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v1_ordinal1(X0) [cnf transformation 60] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
119. ~v1_xboole_0(X0) | v3_ordinal1(X0) [cnf transformation 64] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 65] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
121. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 66] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
128. ~r1_tarski(X0,X1) | X0 = X1 | r2_xboole_0(X0,X1) [cnf transformation 75] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
155. 1 <=> r2_hidden(k1_xboole_0,sK0) [avatar definition]
157. ~r2_hidden(k1_xboole_0,sK0) <- (~1) [avatar component clause 155] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
158. ~1 [avatar split clause 109,155]
160. 2 <=> k1_xboole_0 = sK0 [avatar definition]
162. k1_xboole_0 != sK0 <- (~2) [avatar component clause 160] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
163. ~2 [avatar split clause 108,160]
165. 3 <=> v3_ordinal1(sK0) [avatar definition]
167. v3_ordinal1(sK0) <- (3) [avatar component clause 165] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
168. 3 [avatar split clause 107,165]
175. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
177. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 175] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
178. 5 [avatar split clause 111,175]
279. v3_ordinal1(k1_xboole_0) <- (5) [resolution 119,177] {a:1,w:2,wCS:800,sine:1,thAx:0,allAx:2,thDist:-2}
286. 26 <=> v3_ordinal1(k1_xboole_0) [avatar definition]
288. v3_ordinal1(k1_xboole_0) <- (26) [avatar component clause 286] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
289. 26 | ~5 [avatar split clause 279,175,286]
343. ~r2_hidden(X0,sK0) | ~v1_xboole_0(X0) <- (~1) [superposition 157,120] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
367. v1_ordinal1(k1_xboole_0) <- (26) [resolution 288,121] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
375. 38 <=> v1_ordinal1(k1_xboole_0) [avatar definition]
377. v1_ordinal1(k1_xboole_0) <- (38) [avatar component clause 375] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
378. 38 | ~26 [avatar split clause 367,286,375]
386. r2_xboole_0(k1_xboole_0,X0) | k1_xboole_0 = X0 [resolution 128,113] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
434. k1_xboole_0 = X0 | r2_hidden(k1_xboole_0,X0) | ~v3_ordinal1(X0) | ~v1_ordinal1(k1_xboole_0) [resolution 386,115] {a:2,w:10,wCS:4000,sine:1,thAx:0,allAx:3,thDist:-3}
437. r2_hidden(k1_xboole_0,X0) | k1_xboole_0 = X0 | ~v3_ordinal1(X0) <- (38) [subsumption resolution 434,377] {a:2,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
545. k1_xboole_0 = sK0 | ~v3_ordinal1(sK0) | ~v1_xboole_0(k1_xboole_0) <- (~1, 38) [resolution 437,343] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
554. ~v3_ordinal1(sK0) | ~v1_xboole_0(k1_xboole_0) <- (~1, ~2, 38) [subsumption resolution 545,162] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
555. ~v1_xboole_0(k1_xboole_0) <- (~1, ~2, 3, 38) [subsumption resolution 554,167] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
556. $false <- (~1, ~2, 3, 5, 38) [subsumption resolution 555,177] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
557. 1 | 2 | ~3 | ~5 | ~38 [avatar contradiction clause 556]
559. $false [avatar sat refutation 158,163,168,178,289,378,557] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:210,thDist:-210}
% SZS output end Proof for ordinal3__t8_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 42
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 1
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 47
% Generated clauses: 298
% Active clauses: 117
% Passive clauses: 146
% Final active clauses: 98
% Final passive clauses: 16
% Main loop iterations started: 130
% 
% >>> Simplifying Inferences
% Duplicate literals: 12
% Fw subsumption resolutions: 12
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Equational tautologies: 5
% Forward subsumptions: 80
% Backward subsumptions: 17
% 
% >>> Generating Inferences
% Binary resolution: 144
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 9
% Backward superposition: 14
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 49
% SAT solver unit clauses: 16
% SAT solver binary clauses: 27
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

