i: [1,25,0,10,1,1,0,0] t32_xtuple_0
a: [3,36,0,10,1,25]
s: 36
i: [1,30,0,3,1,0,0,1] reflexivity_r1_tarski
s: 30
i: [1,27,0,5,1,0,0,1] t17_xboole_1
s: 27
i: [1,31,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 31
i: [1,26,0,6,1,0,0,255] d14_xtuple_0
s: 26
i: [1,32,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 32
d: [2,40,1,5,1,0,44,27,32]
s: 40
i: [1,28,0,8,2,0,0,1] t10_xtuple_0
s: 28
d: [2,41,1,5,1,0,44,27,32]
s: 41
i: [1,29,0,11,3,0,0,1] t19_xboole_1
s: 29
d: [2,45,1,14,2,1,40,29,36]
a: [3,53,1,7,1,45]
s: 53
d: [2,59,2,0,0,1,52,27,53,28]
e: 59
d: [2,60,2,5,1,1,40,53,28]
d: [2,62,2,0,0,1,21,60,27]
e: 62
r: 53
a: [3,52,1,7,1,45]
s: 52
a: [3,57,1,7,1,45]
s: 57
d: [2,64,2,0,0,1,52,40,57,28]
e: 64
d: [2,65,2,5,1,1,40,57,28]
d: [2,67,2,0,0,1,21,65,40]
e: 67
f: 67,62,45,25
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t32_xtuple_0
% SZS output start Proof for xtuple_0__t32_xtuple_0
1. ! [X0] : ! [X1] : r1_tarski(k11_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1))) [input t32_xtuple_0]
2. ~! [X0] : ! [X1] : r1_tarski(k11_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1))) [negated conjecture 1]
3. ! [X0,X1] : k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [input commutativity_k3_xboole_0]
10. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1))) [input t10_xtuple_0]
11. ! [X0] : ! [X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [input t17_xboole_1]
12. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k3_xboole_0(X1,X2))) [input t19_xboole_1]
13. ~! [X0,X1] : r1_tarski(k11_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1))) [flattening 2]
14. ! [X0,X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [flattening 11]
15. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1))) [flattening 10]
16. ! [X0,X1,X2] : ((r1_tarski(X0,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,k3_xboole_0(X1,X2))) [flattening 12]
19. ? [X0,X1] : ~r1_tarski(k11_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1))) [ennf transformation 13]
20. ! [X0,X1] : (r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 15]
21. ! [X0,X1,X2] : (r1_tarski(X0,k3_xboole_0(X1,X2)) | (~r1_tarski(X0,X2) | ~r1_tarski(X0,X1))) [ennf transformation 16]
22. ! [X0,X1,X2] : (r1_tarski(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X2) | ~r1_tarski(X0,X1)) [flattening 21]
23. ? [X0,X1] : ~r1_tarski(k11_xtuple_0(k3_xboole_0(X0,X1)),k3_xboole_0(k11_xtuple_0(X0),k11_xtuple_0(X1))) => ~r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1))) [choice axiom]
24. ~r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1))) [skolemisation 19,23]
25. ~r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1))) [cnf transformation 24] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. r1_tarski(k3_xboole_0(X0,X1),X0) [cnf transformation 14] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
28. r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 20] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
29. r1_tarski(X0,k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 22] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
32. k3_xboole_0(X0,X1) = k3_xboole_0(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
34. 1 <=> r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1))) [avatar definition]
36. ~r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k3_xboole_0(k11_xtuple_0(sK0),k11_xtuple_0(sK1))) <- (~1) [avatar component clause 34] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
37. ~1 [avatar split clause 25,34]
40. r1_tarski(k3_xboole_0(X1,X0),X0) [superposition 27,32] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
45. ~r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k11_xtuple_0(sK1)) | ~r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k11_xtuple_0(sK0)) <- (~1) [resolution 29,36] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
51. 2 <=> r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k11_xtuple_0(sK0)) [avatar definition]
53. ~r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k11_xtuple_0(sK0)) <- (~2) [avatar component clause 51] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
55. 3 <=> r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k11_xtuple_0(sK1)) [avatar definition]
57. ~r1_tarski(k11_xtuple_0(k3_xboole_0(sK0,sK1)),k11_xtuple_0(sK1)) <- (~3) [avatar component clause 55] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
58. ~2 | ~3 | 1 [avatar split clause 45,34,55,51]
60. ~r1_tarski(k3_xboole_0(sK0,sK1),sK0) <- (~2) [resolution 53,28] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
62. $false <- (~2) [subsumption resolution 60,27] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
63. 2 [avatar contradiction clause 62]
65. ~r1_tarski(k3_xboole_0(sK0,sK1),sK1) <- (~3) [resolution 57,28] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
67. $false <- (~3) [subsumption resolution 65,40] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
68. 3 [avatar contradiction clause 67]
69. $false [avatar sat refutation 37,58,63,68] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:28,thDist:-28}
% SZS output end Proof for xtuple_0__t32_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 8
% Generated clauses: 30
% Active clauses: 12
% Passive clauses: 14
% Final active clauses: 11
% Final passive clauses: 1
% Main loop iterations started: 13
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 3
% Unit resulting resolution: 2
% Forward superposition: 6
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 8
% SAT solver unit clauses: 5
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

