i: [1,38,0,2,1,1,0,0] t7_orders_2
a: [3,100,0,2,1,38]
s: 100
i: [1,37,0,2,1,1,0,0] t7_orders_2
a: [3,105,0,2,1,37]
s: 105
i: [1,36,0,2,1,1,0,0] t7_orders_2
a: [3,110,0,2,1,36]
s: 110
i: [1,46,0,4,1,1,0,0] t7_orders_2
a: [3,60,0,4,1,46]
s: 60
i: [1,45,0,8,2,1,0,0] t7_orders_2
a: [3,69,0,4,1,45]
s: 69
i: [1,43,0,8,2,1,0,0] t7_orders_2
a: [3,79,0,4,1,43]
s: 79
i: [1,41,0,4,1,1,0,0] t7_orders_2
a: [3,85,0,4,1,41]
s: 85
i: [1,40,0,4,1,1,0,0] t7_orders_2
a: [3,90,0,4,1,40]
s: 90
i: [1,39,0,4,1,1,0,0] t7_orders_2
a: [3,95,0,4,1,39]
s: 95
i: [1,54,0,2,1,0,0,1] existence_l1_orders_2
a: [3,115,0,2,1,54]
s: 115
i: [1,52,0,4,1,0,0,1] existence_m1_subset_1
s: 52
i: [1,55,0,14,4,0,0,2] d6_orders_2
d: [2,56,0,10,3,0,18,55]
s: 56
i: [1,53,0,14,4,0,0,1] irreflexivity_r2_orders_2
s: 53
i: [1,47,0,18,5,0,0,2] d6_orders_2
s: 47
a: [3,65,0,4,1,45]
s: 65
i: [1,49,0,21,6,0,0,2] d6_orders_2
s: 49
i: [1,44,0,8,2,1,0,0] t7_orders_2
a: [3,74,0,4,1,44]
s: 74
i: [1,50,0,23,7,0,0,1] t2_orders_2
s: 50
d: [2,133,1,19,6,1,40,50,79]
d: [2,135,1,15,5,2,21,133,95]
d: [2,136,1,11,4,3,21,135,90]
d: [2,137,1,9,3,4,21,136,100]
d: [2,138,1,7,2,5,21,137,105]
a: [3,142,1,4,1,138]
s: 142
i: [1,51,0,28,8,0,0,1] t3_orders_2
s: 51
d: [2,134,1,19,6,1,40,50,65]
d: [2,144,1,15,5,2,21,134,90]
d: [2,145,1,11,4,3,21,144,85]
d: [2,146,1,9,3,4,21,145,100]
d: [2,147,1,7,2,5,21,146,105]
a: [3,151,1,4,1,147]
s: 151
d: [2,157,1,24,7,1,40,51,79]
d: [2,159,1,20,6,2,21,157,90]
d: [2,160,1,16,5,3,21,159,95]
d: [2,161,1,14,4,4,21,160,100]
d: [2,162,1,12,3,5,21,161,110]
s: 162
d: [2,158,1,24,7,1,40,51,65]
d: [2,163,1,20,6,2,21,158,85]
d: [2,164,1,16,5,3,21,163,90]
d: [2,165,1,14,4,4,21,164,100]
d: [2,166,1,12,3,5,21,165,110]
s: 166
d: [2,167,2,8,2,6,40,166,79]
d: [2,168,2,4,1,7,21,167,95]
a: [3,172,2,4,1,168]
s: 172
r: 162
r: 79
r: 172
i: [1,42,0,8,2,1,0,0] t7_orders_2
d: [2,117,1,14,4,1,40,47,69]
d: [2,118,1,10,3,2,21,117,85]
d: [2,119,1,6,2,3,21,118,90]
d: [2,120,1,4,1,4,21,119,100]
d: [2,122,1,17,5,1,40,49,79]
d: [2,124,1,13,4,2,21,122,90]
d: [2,125,1,9,3,3,21,124,95]
d: [2,126,1,7,2,4,21,125,100]
d: [2,176,3,17,5,1,40,172,49]
d: [2,194,3,13,4,2,21,176,60]
d: [2,195,3,9,3,3,21,194,85]
d: [2,196,3,5,2,4,21,195,95]
d: [2,197,3,3,1,5,21,196,100]
d: [2,132,1,14,4,1,40,74,47]
d: [2,201,1,10,3,2,21,132,90]
d: [2,202,1,6,2,3,21,201,95]
d: [2,203,1,4,1,4,21,202,100]
f: 203,197,168,147,126,120,36,37,38,39,40,41,42,44,45,46
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_2__t7_orders_2
% SZS output start Proof for orders_2__t7_orders_2
1. ! [X0] : ((l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (((r2_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X2)) | (r1_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2))) => r2_orders_2(X0,X1,X3)))))) [input t7_orders_2]
2. ~! [X0] : ((l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (((r2_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X2)) | (r1_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2))) => r2_orders_2(X0,X1,X3)))))) [negated conjecture 1]
3. ! [X0] : (l1_orders_2(X0) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_orders_2(X0,X1,X2) <=> (X1 != X2 & r1_orders_2(X0,X1,X2)))))) [input d6_orders_2]
12. ! [X0] : ((l1_orders_2(X0) & v5_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ((r1_orders_2(X0,X2,X1) & r1_orders_2(X0,X1,X2)) => X1 = X2)))) [input t2_orders_2]
13. ! [X0] : ((l1_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ((r1_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X2)) => r1_orders_2(X0,X1,X3)))))) [input t3_orders_2]
16. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r2_orders_2(X0,X1,X3) & ((r2_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X2)) | (r1_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2)))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0))) [ennf transformation 2]
17. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_orders_2(X0,X1,X3) & ((r2_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X2)) | (r1_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0)) [flattening 16]
18. ! [X0] : (! [X1] : (! [X2] : ((r2_orders_2(X0,X1,X2) <=> (X1 != X2 & r1_orders_2(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0)) [ennf transformation 3]
19. ! [X0] : (! [X1] : (! [X2] : ((X1 = X2 | (~r1_orders_2(X0,X2,X1) | ~r1_orders_2(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v5_orders_2(X0))) [ennf transformation 12]
20. ! [X0] : (! [X1] : (! [X2] : (X1 = X2 | ~r1_orders_2(X0,X2,X1) | ~r1_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v5_orders_2(X0)) [flattening 19]
21. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_orders_2(X0,X1,X3) | (~r1_orders_2(X0,X2,X3) | ~r1_orders_2(X0,X1,X2))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v4_orders_2(X0))) [ennf transformation 13]
22. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (r1_orders_2(X0,X1,X3) | ~r1_orders_2(X0,X2,X3) | ~r1_orders_2(X0,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v4_orders_2(X0)) [flattening 21]
25. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_orders_2(X0,X1,X3) & ((r2_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X2)) | (r1_orders_2(X0,X2,X3) & r2_orders_2(X0,X1,X2))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v5_orders_2(X0) & v4_orders_2(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r2_orders_2(sK0,X1,X3) & ((r2_orders_2(sK0,X2,X3) & r1_orders_2(sK0,X1,X2)) | (r1_orders_2(sK0,X2,X3) & r2_orders_2(sK0,X1,X2))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v5_orders_2(sK0) & v4_orders_2(sK0)) [choice axiom]
26. ? [X1] : (? [X2] : (? [X3] : (~r2_orders_2(sK0,X1,X3) & ((r2_orders_2(sK0,X2,X3) & r1_orders_2(sK0,X1,X2)) | (r1_orders_2(sK0,X2,X3) & r2_orders_2(sK0,X1,X2))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (~r2_orders_2(sK0,sK1,X3) & ((r2_orders_2(sK0,X2,X3) & r1_orders_2(sK0,sK1,X2)) | (r1_orders_2(sK0,X2,X3) & r2_orders_2(sK0,sK1,X2))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
27. ? [X2] : (? [X3] : (~r2_orders_2(sK0,sK1,X3) & ((r2_orders_2(sK0,X2,X3) & r1_orders_2(sK0,sK1,X2)) | (r1_orders_2(sK0,X2,X3) & r2_orders_2(sK0,sK1,X2))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (~r2_orders_2(sK0,sK1,X3) & ((r2_orders_2(sK0,sK2,X3) & r1_orders_2(sK0,sK1,sK2)) | (r1_orders_2(sK0,sK2,X3) & r2_orders_2(sK0,sK1,sK2))) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
28. ? [X3] : (~r2_orders_2(sK0,sK1,X3) & ((r2_orders_2(sK0,sK2,X3) & r1_orders_2(sK0,sK1,sK2)) | (r1_orders_2(sK0,sK2,X3) & r2_orders_2(sK0,sK1,sK2))) & m1_subset_1(X3,u1_struct_0(sK0))) => (~r2_orders_2(sK0,sK1,sK3) & ((r2_orders_2(sK0,sK2,sK3) & r1_orders_2(sK0,sK1,sK2)) | (r1_orders_2(sK0,sK2,sK3) & r2_orders_2(sK0,sK1,sK2))) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
29. (((~r2_orders_2(sK0,sK1,sK3) & ((r2_orders_2(sK0,sK2,sK3) & r1_orders_2(sK0,sK1,sK2)) | (r1_orders_2(sK0,sK2,sK3) & r2_orders_2(sK0,sK1,sK2))) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v5_orders_2(sK0) & v4_orders_2(sK0) [skolemisation 17,28,27,26,25]
30. ! [X0] : (! [X1] : (! [X2] : (((r2_orders_2(X0,X1,X2) | (X1 = X2 | ~r1_orders_2(X0,X1,X2))) & ((X1 != X2 & r1_orders_2(X0,X1,X2)) | ~r2_orders_2(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0)) [nnf transformation 18]
31. ! [X0] : (! [X1] : (! [X2] : (((r2_orders_2(X0,X1,X2) | X1 = X2 | ~r1_orders_2(X0,X1,X2)) & ((X1 != X2 & r1_orders_2(X0,X1,X2)) | ~r2_orders_2(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0)) [flattening 30]
36. v4_orders_2(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
37. v5_orders_2(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
38. l1_orders_2(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
39. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
40. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
41. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
42. r1_orders_2(sK0,sK1,sK2) | r2_orders_2(sK0,sK1,sK2) [cnf transformation 29] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
44. r2_orders_2(sK0,sK2,sK3) | r2_orders_2(sK0,sK1,sK2) [cnf transformation 29] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
45. r2_orders_2(sK0,sK2,sK3) | r1_orders_2(sK0,sK2,sK3) [cnf transformation 29] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
46. ~r2_orders_2(sK0,sK1,sK3) [cnf transformation 29] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
47. ~r2_orders_2(X0,X1,X2) | r1_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 31] {a:0,w:18,wCS:7200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
49. ~r1_orders_2(X0,X1,X2) | X1 = X2 | r2_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) [cnf transformation 31] {a:0,w:21,wCS:8400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
50. ~r1_orders_2(X0,X2,X1) | X1 = X2 | ~r1_orders_2(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v5_orders_2(X0) [cnf transformation 20] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. ~r1_orders_2(X0,X2,X3) | r1_orders_2(X0,X1,X3) | ~r1_orders_2(X0,X1,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v4_orders_2(X0) [cnf transformation 22] {a:0,w:28,wCS:11200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. 1 <=> r2_orders_2(sK0,sK1,sK3) [avatar definition]
60. ~r2_orders_2(sK0,sK1,sK3) <- (~1) [avatar component clause 58] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
61. ~1 [avatar split clause 46,58]
63. 2 <=> r1_orders_2(sK0,sK2,sK3) [avatar definition]
65. r1_orders_2(sK0,sK2,sK3) <- (2) [avatar component clause 63] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
67. 3 <=> r2_orders_2(sK0,sK2,sK3) [avatar definition]
69. r2_orders_2(sK0,sK2,sK3) <- (3) [avatar component clause 67] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
70. 2 | 3 [avatar split clause 45,67,63]
72. 4 <=> r2_orders_2(sK0,sK1,sK2) [avatar definition]
74. r2_orders_2(sK0,sK1,sK2) <- (4) [avatar component clause 72] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
75. 4 | 3 [avatar split clause 44,67,72]
77. 5 <=> r1_orders_2(sK0,sK1,sK2) [avatar definition]
79. r1_orders_2(sK0,sK1,sK2) <- (5) [avatar component clause 77] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
81. 4 | 5 [avatar split clause 42,77,72]
83. 6 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
85. m1_subset_1(sK3,u1_struct_0(sK0)) <- (6) [avatar component clause 83] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
86. 6 [avatar split clause 41,83]
88. 7 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
90. m1_subset_1(sK2,u1_struct_0(sK0)) <- (7) [avatar component clause 88] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
91. 7 [avatar split clause 40,88]
93. 8 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
95. m1_subset_1(sK1,u1_struct_0(sK0)) <- (8) [avatar component clause 93] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
96. 8 [avatar split clause 39,93]
98. 9 <=> l1_orders_2(sK0) [avatar definition]
100. l1_orders_2(sK0) <- (9) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
101. 9 [avatar split clause 38,98]
103. 10 <=> v5_orders_2(sK0) [avatar definition]
105. v5_orders_2(sK0) <- (10) [avatar component clause 103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
106. 10 [avatar split clause 37,103]
108. 11 <=> v4_orders_2(sK0) [avatar definition]
110. v4_orders_2(sK0) <- (11) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
111. 11 [avatar split clause 36,108]
117. r1_orders_2(sK0,sK2,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (3) [resolution 47,69] {a:1,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
118. r1_orders_2(sK0,sK2,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (3, 6) [subsumption resolution 117,85] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
119. r1_orders_2(sK0,sK2,sK3) | ~l1_orders_2(sK0) <- (3, 6, 7) [subsumption resolution 118,90] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
120. r1_orders_2(sK0,sK2,sK3) <- (3, 6, 7, 9) [subsumption resolution 119,100] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
121. 2 | ~3 | ~6 | ~7 | ~9 [avatar split clause 120,98,88,83,67,63]
122. sK1 = sK2 | r2_orders_2(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (5) [resolution 49,79] {a:1,w:17,wCS:1700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
124. sK1 = sK2 | r2_orders_2(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (5, 7) [subsumption resolution 122,90] {a:1,w:13,wCS:1300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
125. sK1 = sK2 | r2_orders_2(sK0,sK1,sK2) | ~l1_orders_2(sK0) <- (5, 7, 8) [subsumption resolution 124,95] {a:1,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
126. sK1 = sK2 | r2_orders_2(sK0,sK1,sK2) <- (5, 7, 8, 9) [subsumption resolution 125,100] {a:1,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
128. 13 <=> sK1 = sK2 [avatar definition]
131. 4 | 13 | ~5 | ~7 | ~8 | ~9 [avatar split clause 126,98,93,88,77,128,72]
132. r1_orders_2(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (4) [resolution 74,47] {a:1,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
134. sK2 = sK3 | ~r1_orders_2(sK0,sK3,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v5_orders_2(sK0) <- (2) [resolution 50,65] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
144. sK2 = sK3 | ~r1_orders_2(sK0,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v5_orders_2(sK0) <- (2, 7) [subsumption resolution 134,90] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
145. sK2 = sK3 | ~r1_orders_2(sK0,sK3,sK2) | ~l1_orders_2(sK0) | ~v5_orders_2(sK0) <- (2, 6, 7) [subsumption resolution 144,85] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
146. sK2 = sK3 | ~r1_orders_2(sK0,sK3,sK2) | ~v5_orders_2(sK0) <- (2, 6, 7, 9) [subsumption resolution 145,100] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
147. sK2 = sK3 | ~r1_orders_2(sK0,sK3,sK2) <- (2, 6, 7, 9, 10) [subsumption resolution 146,105] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
149. 15 <=> r1_orders_2(sK0,sK3,sK2) [avatar definition]
153. 16 <=> sK2 = sK3 [avatar definition]
156. ~15 | 16 | ~2 | ~6 | ~7 | ~9 | ~10 [avatar split clause 147,103,98,88,83,63,153,149]
158. r1_orders_2(sK0,X1,sK3) | ~r1_orders_2(sK0,X1,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v4_orders_2(sK0) <- (2) [resolution 51,65] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
163. r1_orders_2(sK0,X1,sK3) | ~r1_orders_2(sK0,X1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v4_orders_2(sK0) <- (2, 6) [subsumption resolution 158,85] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
164. r1_orders_2(sK0,X1,sK3) | ~r1_orders_2(sK0,X1,sK2) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v4_orders_2(sK0) <- (2, 6, 7) [subsumption resolution 163,90] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
165. r1_orders_2(sK0,X1,sK3) | ~r1_orders_2(sK0,X1,sK2) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~v4_orders_2(sK0) <- (2, 6, 7, 9) [subsumption resolution 164,100] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
166. ~r1_orders_2(sK0,X1,sK2) | r1_orders_2(sK0,X1,sK3) | ~m1_subset_1(X1,u1_struct_0(sK0)) <- (2, 6, 7, 9, 11) [subsumption resolution 165,110] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
167. r1_orders_2(sK0,sK1,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) <- (2, 5, 6, 7, 9, 11) [resolution 166,79] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
168. r1_orders_2(sK0,sK1,sK3) <- (2, 5, 6, 7, 8, 9, 11) [subsumption resolution 167,95] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
170. 17 <=> r1_orders_2(sK0,sK1,sK3) [avatar definition]
172. r1_orders_2(sK0,sK1,sK3) <- (17) [avatar component clause 170] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
173. 17 | ~2 | ~5 | ~6 | ~7 | ~8 | ~9 | ~11 [avatar split clause 168,108,98,93,88,83,77,63,170]
176. sK1 = sK3 | r2_orders_2(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (17) [resolution 172,49] {a:3,w:17,wCS:1700,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
190. 19 <=> sK1 = sK3 [avatar definition]
194. sK1 = sK3 | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (~1, 17) [subsumption resolution 176,60] {a:3,w:13,wCS:1300,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
195. sK1 = sK3 | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (~1, 6, 17) [subsumption resolution 194,85] {a:3,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
196. sK1 = sK3 | ~l1_orders_2(sK0) <- (~1, 6, 8, 17) [subsumption resolution 195,95] {a:3,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
197. sK1 = sK3 <- (~1, 6, 8, 9, 17) [subsumption resolution 196,100] {a:3,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:12,thDist:-12}
198. 19 | 1 | ~6 | ~8 | ~9 | ~17 [avatar split clause 197,170,98,93,83,58,190]
199. sK1 != sK3 | r1_orders_2(sK0,sK3,sK2) | ~r1_orders_2(sK0,sK1,sK2) [theory tautology sat conflict] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
200. sK2 != sK3 | r2_orders_2(sK0,sK1,sK3) | ~r2_orders_2(sK0,sK1,sK2) [theory tautology sat conflict] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
201. r1_orders_2(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) <- (4, 7) [subsumption resolution 132,90] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
202. r1_orders_2(sK0,sK1,sK2) | ~l1_orders_2(sK0) <- (4, 7, 8) [subsumption resolution 201,95] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
203. r1_orders_2(sK0,sK1,sK2) <- (4, 7, 8, 9) [subsumption resolution 202,100] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
204. 5 | ~4 | ~7 | ~8 | ~9 [avatar split clause 203,98,93,88,72,77]
205. sK1 != sK2 | r2_orders_2(sK0,sK1,sK3) | ~r2_orders_2(sK0,sK2,sK3) [theory tautology sat conflict] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
206. $false [avatar sat refutation 61,70,75,81,86,91,96,101,106,111,121,131,156,173,198,199,200,204,205] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:192,thDist:-192}
% SZS output end Proof for orders_2__t7_orders_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 88
% Active clauses: 23
% Passive clauses: 27
% Final active clauses: 20
% Final passive clauses: 2
% Main loop iterations started: 24
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 38
% 
% >>> Deletion Inferences
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 12
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 16
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 25
% SAT solver unit clauses: 8
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

