i: [1,79,0,2,1,1,0,0] t16_boolealg
a: [3,123,0,2,1,79]
s: 123
i: [1,78,0,2,1,1,0,0] t16_boolealg
a: [3,128,0,2,1,78]
s: 128
i: [1,77,0,2,1,1,0,0] t16_boolealg
a: [3,133,0,2,1,77]
s: 133
i: [1,76,0,2,1,1,0,0] t16_boolealg
a: [3,138,0,2,1,76]
s: 138
i: [1,80,0,4,1,1,0,0] t16_boolealg
a: [3,118,0,4,1,80]
s: 118
i: [1,81,0,5,1,1,0,0] t16_boolealg
a: [3,113,0,5,1,81]
s: 113
i: [1,108,0,2,1,0,0,2] existence_l1_lattices
a: [3,143,0,2,1,108]
s: 143
i: [1,109,0,2,1,0,0,1] existence_l3_lattices
a: [3,148,0,2,1,109]
s: 148
i: [1,82,0,4,2,0,0,255] dt_l3_lattices
s: 82
i: [1,95,0,4,1,0,0,1] existence_m1_subset_1
s: 95
i: [1,86,0,8,4,0,0,255] cc1_lattices
s: 86
d: [2,150,1,2,1,1,40,82,123]
a: [3,155,1,2,1,150]
s: 155
i: [1,87,0,8,4,0,0,255] cc1_lattices
s: 87
d: [2,162,1,6,3,1,40,86,123]
d: [2,164,1,4,2,2,21,162,133]
d: [2,165,1,2,1,3,21,164,138]
a: [3,169,1,2,1,165]
s: 169
i: [1,88,0,8,4,0,0,255] cc1_lattices
s: 88
i: [1,89,0,8,4,0,0,255] cc1_lattices
s: 89
d: [2,184,1,6,3,1,40,87,123]
d: [2,186,1,4,2,2,21,184,133]
d: [2,187,1,2,1,3,21,186,138]
a: [3,191,1,2,1,187]
s: 191
i: [1,90,0,8,4,0,0,255] cc1_lattices
s: 90
d: [2,198,1,6,3,1,40,88,123]
d: [2,200,1,4,2,2,21,198,133]
d: [2,201,1,2,1,3,21,200,138]
a: [3,205,1,2,1,201]
s: 205
i: [1,91,0,8,4,0,0,255] cc1_lattices
s: 91
i: [1,92,0,9,3,0,0,1] dt_k5_lattices
s: 92
d: [2,212,1,6,3,1,40,89,123]
d: [2,214,1,4,2,2,21,212,133]
d: [2,215,1,2,1,3,21,214,138]
a: [3,219,1,2,1,215]
s: 219
i: [1,84,0,18,9,0,0,255] cc2_lattices
s: 84
d: [2,226,1,6,3,1,40,90,123]
d: [2,228,1,4,2,2,21,226,133]
d: [2,229,1,2,1,3,21,228,138]
a: [3,233,1,2,1,229]
s: 233
i: [1,101,0,19,5,0,0,255] dt_k2_lattices
s: 101
i: [1,96,0,20,7,0,0,255] rd2_lattices
s: 96
d: [2,240,1,6,3,1,40,91,123]
d: [2,242,1,4,2,2,21,240,133]
d: [2,243,1,2,1,3,21,242,138]
a: [3,247,1,2,1,243]
s: 247
i: [1,97,0,20,6,0,0,1] rd4_lattices
s: 97
d: [2,151,1,2,1,1,40,82,148]
a: [3,160,1,2,1,151]
s: 160
i: [1,98,0,21,6,0,0,2] dt_k4_lattices
s: 98
i: [1,102,0,22,7,0,0,255] symmetry_r3_boolealg
s: 102
d: [2,163,1,6,3,1,40,86,148]
a: [3,182,1,2,1,163]
s: 182
i: [1,103,0,22,7,0,0,1] symmetry_r4_boolealg
s: 103
d: [2,259,1,20,6,1,40,103,113]
d: [2,260,1,16,5,2,21,259,118]
d: [2,261,1,14,4,3,21,260,123]
d: [2,262,1,12,3,4,21,261,133]
d: [2,263,1,10,2,5,21,262,138]
a: [3,267,1,5,1,263]
s: 267
d: [2,273,2,0,0,3,52,138,155,267,92]
e: 273
d: [2,274,2,4,2,1,40,267,92]
d: [2,276,2,2,1,2,21,274,155]
d: [2,277,2,0,0,3,21,276,138]
e: 277
r: 267
i: [1,104,0,22,7,0,0,255] redefinition_r3_boolealg
s: 104
i: [1,105,0,22,7,0,0,255] redefinition_r3_boolealg
s: 105
a: [3,266,1,5,1,263]
s: 266
i: [1,106,0,22,7,0,0,1] redefinition_r4_boolealg
s: 106
a: [3,271,1,5,1,263]
s: 271
i: [1,107,0,22,7,0,0,1] redefinition_r4_boolealg
s: 107
i: [1,99,0,23,6,0,0,255] redefinition_k4_lattices
s: 99
d: [2,279,1,20,6,1,40,106,113]
d: [2,280,1,15,5,2,21,279,266]
d: [2,281,1,11,4,3,21,280,118]
d: [2,282,1,9,3,4,21,281,123]
d: [2,283,1,7,2,5,21,282,133]
d: [2,284,1,5,1,6,21,283,138]
a: [3,288,1,5,1,284]
s: 288
i: [1,100,0,23,6,0,0,2] commutativity_k4_lattices
s: 100
d: [2,290,2,20,6,1,40,271,106]
d: [2,292,2,16,5,2,21,290,118]
d: [2,293,2,11,4,3,21,292,266]
d: [2,294,2,9,3,4,21,293,123]
d: [2,295,2,7,2,5,21,294,133]
d: [2,296,2,5,1,6,21,295,138]
a: [3,300,2,5,1,296]
s: 300
i: [1,93,0,25,7,0,0,2] d4_boolealg
s: 93
i: [1,94,0,25,7,0,0,2] d4_boolealg
s: 94
d: [2,305,2,20,6,1,40,288,105]
d: [2,306,2,15,5,2,21,305,266]
d: [2,307,2,11,4,3,21,306,118]
d: [2,308,2,9,3,4,21,307,123]
d: [2,309,2,7,2,5,21,308,133]
d: [2,310,2,5,1,6,21,309,138]
a: [3,314,2,5,1,310]
s: 314
d: [2,348,1,36,13,0,44,93,96]
d: [2,356,1,24,9,0,18,348]
d: [2,357,1,22,8,0,21,356,88]
d: [2,358,1,20,7,0,21,357,90]
d: [2,359,1,18,6,0,21,358,91]
s: 359
d: [2,337,3,20,6,1,40,300,105]
d: [2,338,3,16,5,2,21,337,118]
d: [2,339,3,11,4,3,21,338,266]
d: [2,340,3,9,3,4,21,339,123]
d: [2,341,3,7,2,5,21,340,133]
d: [2,342,3,5,1,6,21,341,138]
a: [3,346,3,5,1,342]
s: 346
d: [2,364,1,36,13,0,44,94,96]
d: [2,379,1,24,9,0,18,364]
d: [2,380,1,22,8,0,21,379,88]
d: [2,381,1,20,7,0,21,380,90]
d: [2,382,1,18,6,0,21,381,91]
s: 382
d: [2,367,1,36,13,0,44,96,94]
d: [2,376,1,24,9,0,18,367]
d: [2,387,1,22,8,0,21,376,88]
d: [2,388,1,20,7,0,21,387,90]
d: [2,389,1,18,6,0,21,388,91]
s: 389
d: [2,409,2,14,5,1,40,382,118]
d: [2,417,2,12,4,2,21,409,123]
d: [2,418,2,10,3,3,21,417,133]
d: [2,419,2,8,2,4,21,418,138]
a: [3,423,2,4,1,419]
s: 423
r: 113
r: 266
r: 271
r: 288
r: 300
r: 314
r: 346
r: 423
s: 113
a: [3,422,2,4,1,419]
s: 422
s: 266
s: 271
a: [3,427,2,4,1,419]
s: 427
d: [2,429,0,4,1,2,23,113,423]
a: [3,442,0,4,1,429]
s: 442
d: [2,434,2,4,1,2,23,314,423]
a: [3,450,2,4,1,434]
s: 450
s: 288
d: [2,349,1,37,12,0,44,93,97]
d: [2,354,1,27,8,0,18,349]
d: [2,355,1,22,7,0,19,354]
s: 355
s: 300
d: [2,511,3,0,0,7,52,128,138,133,123,118,266,300,355]
e: 511
d: [2,512,3,17,6,1,40,300,355]
d: [2,516,3,13,5,2,21,512,118]
d: [2,517,3,8,4,3,21,516,266]
d: [2,518,3,6,3,4,21,517,123]
d: [2,519,3,4,2,5,21,518,133]
d: [2,520,3,2,1,6,21,519,138]
d: [2,521,3,0,0,7,21,520,128]
e: 521
d: [2,461,1,20,6,1,40,113,103]
d: [2,467,1,16,5,2,21,461,118]
d: [2,468,1,14,4,3,21,467,123]
d: [2,469,1,12,3,4,21,468,133]
d: [2,470,1,10,2,5,21,469,138]
d: [2,473,2,20,6,1,40,271,106]
d: [2,475,2,16,5,2,21,473,118]
d: [2,476,2,11,4,3,21,475,266]
d: [2,477,2,9,3,4,21,476,123]
d: [2,478,2,7,2,5,21,477,133]
d: [2,479,2,5,1,6,21,478,138]
f: 521,479,470,277,150,76,77,78,79,80,81
% Refutation found. Thanks to Tanya!
% SZS status Theorem for boolealg__t16_boolealg
% SZS output start Proof for boolealg__t16_boolealg
1. ! [X0] : ((l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ~r4_boolealg(X0,X1,k5_lattices(X0)))) [input t16_boolealg]
2. ~! [X0] : ((l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ~r4_boolealg(X0,X1,k5_lattices(X0)))) [negated conjecture 1]
6. ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_boolealg(X0,X1,X2) <=> k5_lattices(X0) != k4_lattices(X0,X1,X2))))) [input d4_boolealg]
9. ! [X0] : ((l1_lattices(X0) & ~v2_struct_0(X0)) => m1_subset_1(k5_lattices(X0),u1_struct_0(X0))) [input dt_k5_lattices]
13. ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))) [input dt_l3_lattices]
22. ! [X0,X1] : ((m1_subset_1(X1,u1_struct_0(X0)) & l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => k5_lattices(X0) = k4_lattices(X0,k5_lattices(X0),X1)) [input rd4_lattices]
25. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (r4_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2))) [input redefinition_r4_boolealg]
27. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (r4_boolealg(X0,X1,X2) => r4_boolealg(X0,X2,X1))) [input symmetry_r4_boolealg]
31. ! [X0] : (l3_lattices(X0) => l1_lattices(X0)) [pure predicate removal 13]
33. ? [X0] : (? [X1] : (r4_boolealg(X0,X1,k5_lattices(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & (l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
34. ? [X0] : (? [X1] : (r4_boolealg(X0,X1,k5_lattices(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) [flattening 33]
35. ! [X0] : (l1_lattices(X0) | ~l3_lattices(X0)) [ennf transformation 31]
40. ! [X0] : (m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | (~l1_lattices(X0) | v2_struct_0(X0))) [ennf transformation 9]
41. ! [X0] : (m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | ~l1_lattices(X0) | v2_struct_0(X0)) [flattening 40]
42. ! [X0] : (! [X1] : (! [X2] : ((r2_boolealg(X0,X1,X2) <=> k5_lattices(X0) != k4_lattices(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 6]
43. ! [X0] : (! [X1] : (! [X2] : ((r2_boolealg(X0,X1,X2) <=> k5_lattices(X0) != k4_lattices(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 42]
46. ! [X0,X1] : (k5_lattices(X0) = k4_lattices(X0,k5_lattices(X0),X1) | (~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 22]
47. ! [X0,X1] : (k5_lattices(X0) = k4_lattices(X0,k5_lattices(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 46]
58. ! [X0,X1,X2] : ((r4_boolealg(X0,X2,X1) | ~r4_boolealg(X0,X1,X2)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 27]
59. ! [X0,X1,X2] : (r4_boolealg(X0,X2,X1) | ~r4_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 58]
62. ! [X0,X1,X2] : ((r4_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0))) [ennf transformation 25]
63. ! [X0,X1,X2] : ((r4_boolealg(X0,X1,X2) <=> r2_boolealg(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [flattening 62]
64. ? [X0] : (? [X1] : (r4_boolealg(X0,X1,k5_lattices(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~v2_struct_0(X0)) => (? [X1] : (r4_boolealg(sK0,X1,k5_lattices(sK0)) & m1_subset_1(X1,u1_struct_0(sK0))) & l3_lattices(sK0) & v13_lattices(sK0) & v10_lattices(sK0) & ~v2_struct_0(sK0)) [choice axiom]
65. ? [X1] : (r4_boolealg(sK0,X1,k5_lattices(sK0)) & m1_subset_1(X1,u1_struct_0(sK0))) => (r4_boolealg(sK0,sK1,k5_lattices(sK0)) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
66. (r4_boolealg(sK0,sK1,k5_lattices(sK0)) & m1_subset_1(sK1,u1_struct_0(sK0))) & l3_lattices(sK0) & v13_lattices(sK0) & v10_lattices(sK0) & ~v2_struct_0(sK0) [skolemisation 34,65,64]
67. ! [X0] : (! [X1] : (! [X2] : (((r2_boolealg(X0,X1,X2) | k5_lattices(X0) = k4_lattices(X0,X1,X2)) & (k5_lattices(X0) != k4_lattices(X0,X1,X2) | ~r2_boolealg(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [nnf transformation 43]
71. ! [X0,X1,X2] : (((r4_boolealg(X0,X1,X2) | ~r2_boolealg(X0,X1,X2)) & (r2_boolealg(X0,X1,X2) | ~r4_boolealg(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0)) [nnf transformation 63]
76. ~v2_struct_0(sK0) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
77. v10_lattices(sK0) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
78. v13_lattices(sK0) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
79. l3_lattices(sK0) [cnf transformation 66] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
80. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 66] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
81. r4_boolealg(sK0,sK1,k5_lattices(sK0)) [cnf transformation 66] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
82. ~l3_lattices(X0) | l1_lattices(X0) [cnf transformation 35] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
92. m1_subset_1(k5_lattices(X0),u1_struct_0(X0)) | ~l1_lattices(X0) | v2_struct_0(X0) [cnf transformation 41] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
93. k5_lattices(X0) != k4_lattices(X0,X1,X2) | ~r2_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 67] {a:0,w:25,wCS:10000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
97. k5_lattices(X0) = k4_lattices(X0,k5_lattices(X0),X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v13_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 47] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
103. ~r4_boolealg(X0,X1,X2) | r4_boolealg(X0,X2,X1) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 59] {a:0,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
106. ~r4_boolealg(X0,X1,X2) | r2_boolealg(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l3_lattices(X0) | ~v10_lattices(X0) | v2_struct_0(X0) [cnf transformation 71] {a:0,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
111. 1 <=> r4_boolealg(sK0,sK1,k5_lattices(sK0)) [avatar definition]
113. r4_boolealg(sK0,sK1,k5_lattices(sK0)) <- (1) [avatar component clause 111] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 1 [avatar split clause 81,111]
116. 2 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
118. m1_subset_1(sK1,u1_struct_0(sK0)) <- (2) [avatar component clause 116] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
119. 2 [avatar split clause 80,116]
121. 3 <=> l3_lattices(sK0) [avatar definition]
123. l3_lattices(sK0) <- (3) [avatar component clause 121] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. 3 [avatar split clause 79,121]
126. 4 <=> v13_lattices(sK0) [avatar definition]
128. v13_lattices(sK0) <- (4) [avatar component clause 126] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. 4 [avatar split clause 78,126]
131. 5 <=> v10_lattices(sK0) [avatar definition]
133. v10_lattices(sK0) <- (5) [avatar component clause 131] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. 5 [avatar split clause 77,131]
136. 6 <=> v2_struct_0(sK0) [avatar definition]
138. ~v2_struct_0(sK0) <- (~6) [avatar component clause 136] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
139. ~6 [avatar split clause 76,136]
150. l1_lattices(sK0) <- (3) [resolution 82,123] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
153. 9 <=> l1_lattices(sK0) [avatar definition]
155. l1_lattices(sK0) <- (9) [avatar component clause 153] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
156. 9 | ~3 [avatar split clause 150,121,153]
265. 25 <=> m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) [avatar definition]
266. m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) <- (25) [avatar component clause 265] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
267. ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) <- (~25) [avatar component clause 265] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
269. 26 <=> r4_boolealg(sK0,k5_lattices(sK0),sK1) [avatar definition]
271. r4_boolealg(sK0,k5_lattices(sK0),sK1) <- (26) [avatar component clause 269] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
274. ~l1_lattices(sK0) | v2_struct_0(sK0) <- (~25) [resolution 267,92] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
276. v2_struct_0(sK0) <- (9, ~25) [subsumption resolution 274,155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
277. $false <- (~6, 9, ~25) [subsumption resolution 276,138] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
278. 6 | ~9 | 25 [avatar contradiction clause 277]
298. 28 <=> r2_boolealg(sK0,k5_lattices(sK0),sK1) [avatar definition]
300. r2_boolealg(sK0,k5_lattices(sK0),sK1) <- (28) [avatar component clause 298] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
349. k5_lattices(X2) != k5_lattices(X2) | ~r2_boolealg(X2,k5_lattices(X2),X3) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~m1_subset_1(k5_lattices(X2),u1_struct_0(X2)) | ~l3_lattices(X2) | ~v10_lattices(X2) | v2_struct_0(X2) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~l3_lattices(X2) | ~v13_lattices(X2) | ~v10_lattices(X2) | v2_struct_0(X2) [superposition 93,97] {a:1,w:37,wCS:14800,sine:1,thAx:0,allAx:2,thDist:-2}
354. k5_lattices(X2) != k5_lattices(X2) | ~r2_boolealg(X2,k5_lattices(X2),X3) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~m1_subset_1(k5_lattices(X2),u1_struct_0(X2)) | ~l3_lattices(X2) | ~v10_lattices(X2) | v2_struct_0(X2) | ~v13_lattices(X2) [duplicate literal removal 349] {a:1,w:27,wCS:10800,sine:1,thAx:0,allAx:2,thDist:-2}
355. ~r2_boolealg(X2,k5_lattices(X2),X3) | ~m1_subset_1(X3,u1_struct_0(X2)) | ~m1_subset_1(k5_lattices(X2),u1_struct_0(X2)) | ~l3_lattices(X2) | ~v10_lattices(X2) | v2_struct_0(X2) | ~v13_lattices(X2) [trivial inequality removal 354] {a:1,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
461. r4_boolealg(sK0,k5_lattices(sK0),sK1) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1) [resolution 113,103] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
467. r4_boolealg(sK0,k5_lattices(sK0),sK1) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1, 2) [subsumption resolution 461,118] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
468. r4_boolealg(sK0,k5_lattices(sK0),sK1) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (1, 2, 3) [subsumption resolution 467,123] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
469. r4_boolealg(sK0,k5_lattices(sK0),sK1) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | v2_struct_0(sK0) <- (1, 2, 3, 5) [subsumption resolution 468,133] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
470. r4_boolealg(sK0,k5_lattices(sK0),sK1) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) <- (1, 2, 3, 5, ~6) [subsumption resolution 469,138] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
471. ~25 | 26 | ~1 | ~2 | ~3 | ~5 | 6 [avatar split clause 470,136,131,121,116,111,269,265]
473. r2_boolealg(sK0,k5_lattices(sK0),sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (26) [resolution 271,106] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
475. r2_boolealg(sK0,k5_lattices(sK0),sK1) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (2, 26) [subsumption resolution 473,118] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
476. r2_boolealg(sK0,k5_lattices(sK0),sK1) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (2, 25, 26) [subsumption resolution 475,266] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
477. r2_boolealg(sK0,k5_lattices(sK0),sK1) | ~v10_lattices(sK0) | v2_struct_0(sK0) <- (2, 3, 25, 26) [subsumption resolution 476,123] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
478. r2_boolealg(sK0,k5_lattices(sK0),sK1) | v2_struct_0(sK0) <- (2, 3, 5, 25, 26) [subsumption resolution 477,133] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
479. r2_boolealg(sK0,k5_lattices(sK0),sK1) <- (2, 3, 5, ~6, 25, 26) [subsumption resolution 478,138] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
480. 28 | ~2 | ~3 | ~5 | 6 | ~25 | ~26 [avatar split clause 479,269,265,136,131,121,116,298]
512. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) | ~v13_lattices(sK0) <- (28) [resolution 300,355] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
516. ~m1_subset_1(k5_lattices(sK0),u1_struct_0(sK0)) | ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) | ~v13_lattices(sK0) <- (2, 28) [subsumption resolution 512,118] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
517. ~l3_lattices(sK0) | ~v10_lattices(sK0) | v2_struct_0(sK0) | ~v13_lattices(sK0) <- (2, 25, 28) [subsumption resolution 516,266] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
518. ~v10_lattices(sK0) | v2_struct_0(sK0) | ~v13_lattices(sK0) <- (2, 3, 25, 28) [subsumption resolution 517,123] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
519. v2_struct_0(sK0) | ~v13_lattices(sK0) <- (2, 3, 5, 25, 28) [subsumption resolution 518,133] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
520. ~v13_lattices(sK0) <- (2, 3, 5, ~6, 25, 28) [subsumption resolution 519,138] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
521. $false <- (2, 3, 4, 5, ~6, 25, 28) [subsumption resolution 520,128] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
522. ~2 | ~3 | ~4 | ~5 | 6 | ~25 | ~28 [avatar contradiction clause 521]
528. $false [avatar sat refutation 114,119,124,129,134,139,156,278,471,480,522] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:379,thDist:-379}
% SZS output end Proof for boolealg__t16_boolealg
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 34
% Generated clauses: 307
% Active clauses: 61
% Passive clauses: 80
% Final active clauses: 52
% Final passive clauses: 16
% Main loop iterations started: 62
% 
% >>> Simplifying Inferences
% Duplicate literals: 79
% Trivial inequalities: 1
% Fw subsumption resolutions: 122
% Bw demodulations: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Equational tautologies: 3
% Forward subsumptions: 21
% 
% >>> Generating Inferences
% Binary resolution: 50
% Unit resulting resolution: 2
% Forward superposition: 11
% Backward superposition: 13
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 10
% Split components: 26
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 47
% SAT solver unit clauses: 8
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

