i: [1,29,0,11,4,1,0,0] t14_gate_1
a: [3,43,0,2,1,29]
s: 43
a: [3,38,0,5,1,29]
s: 38
i: [1,34,0,2,1,0,0,1] rc2_xboole_0
a: [3,59,0,2,1,34]
s: 59
i: [1,35,0,2,1,0,0,1] rc1_xboole_0
a: [3,64,0,2,1,35]
s: 64
i: [1,30,0,7,2,0,0,1] fc20_gate_1
s: 30
a: [3,50,0,2,1,29]
s: 50
i: [1,31,0,7,2,0,0,1] fc19_gate_1
s: 31
a: [3,46,0,2,1,29]
s: 46
i: [1,32,0,7,2,0,0,1] fc18_gate_1
s: 32
d: [2,70,1,0,0,2,52,43,38,32]
e: 70
d: [2,71,1,2,1,1,40,32,38]
d: [2,73,1,0,0,2,21,71,43]
e: 73
r: 38
r: 43
a: [3,42,0,2,1,29]
s: 42
a: [3,39,0,5,1,29]
s: 39
i: [1,33,0,11,4,0,0,1] fc17_gate_1
s: 33
d: [2,75,1,0,0,4,52,42,46,50,39,33]
e: 75
d: [2,79,1,6,3,1,40,33,39]
d: [2,81,1,4,2,2,21,79,50]
d: [2,82,1,2,1,3,21,81,46]
d: [2,83,1,0,0,4,21,82,42]
e: 83
i: [1,28,0,7,2,1,0,0] t14_gate_1
i: [1,27,0,7,2,1,0,0] t14_gate_1
i: [1,26,0,7,2,1,0,0] t14_gate_1
d: [2,66,1,2,1,1,40,30,38]
d: [2,68,1,2,1,1,40,31,38]
f: 83,73,68,66,26,27,28,29
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t14_gate_1
% SZS output start Proof for gate_1__t14_gate_1
1. ! [X0] : ! [X1] : ! [X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) <=> (~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input t14_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) <=> (~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [negated conjecture 1]
5. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc17_gate_1]
6. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc18_gate_1]
7. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X0,X2))) [input fc19_gate_1]
8. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X2,X0))) [input fc20_gate_1]
11. ~! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) <=> (~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 2]
12. ? [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) <~> (~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [ennf transformation 11]
13. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
14. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0)) [ennf transformation 7]
15. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0)) [ennf transformation 6]
16. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 5]
17. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 16]
18. ? [X0,X1,X2] : (((v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) | v1_xboole_0(k8_gate_1(X0,X1,X2))) & ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k8_gate_1(X0,X1,X2)))) [nnf transformation 12]
19. ? [X0,X1,X2] : ((v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(X0,X1,X2))) & ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k8_gate_1(X0,X1,X2)))) [flattening 18]
20. ? [X0,X1,X2] : ((v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k8_gate_1(X0,X1,X2))) & ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k8_gate_1(X0,X1,X2)))) => ((v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(sK0,sK1,sK2))) & ((~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k8_gate_1(sK0,sK1,sK2)))) [choice axiom]
21. (v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(sK0,sK1,sK2))) & ((~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k8_gate_1(sK0,sK1,sK2))) [skolemisation 19,20]
26. ~v1_xboole_0(sK0) | ~v1_xboole_0(k8_gate_1(sK0,sK1,sK2)) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. ~v1_xboole_0(sK1) | ~v1_xboole_0(k8_gate_1(sK0,sK1,sK2)) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. ~v1_xboole_0(sK2) | ~v1_xboole_0(k8_gate_1(sK0,sK1,sK2)) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k8_gate_1(sK0,sK1,sK2)) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0) [cnf transformation 13] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
31. v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0) [cnf transformation 14] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
32. v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0) [cnf transformation 15] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. ~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 17] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. 1 <=> v1_xboole_0(k8_gate_1(sK0,sK1,sK2)) [avatar definition]
38. ~v1_xboole_0(k8_gate_1(sK0,sK1,sK2)) <- (~1) [avatar component clause 37] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
39. v1_xboole_0(k8_gate_1(sK0,sK1,sK2)) <- (1) [avatar component clause 37] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
41. 2 <=> v1_xboole_0(sK0) [avatar definition]
42. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 41] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
43. v1_xboole_0(sK0) <- (2) [avatar component clause 41] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
45. 3 <=> v1_xboole_0(sK1) [avatar definition]
46. ~v1_xboole_0(sK1) <- (~3) [avatar component clause 45] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
49. 4 <=> v1_xboole_0(sK2) [avatar definition]
50. ~v1_xboole_0(sK2) <- (~4) [avatar component clause 49] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
52. 1 | 2 | 3 | 4 [avatar split clause 29,49,45,41,37]
53. ~1 | ~4 [avatar split clause 28,49,37]
54. ~1 | ~3 [avatar split clause 27,45,37]
55. ~1 | ~2 [avatar split clause 26,41,37]
66. ~v1_xboole_0(sK2) <- (~1) [resolution 30,38] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
67. ~4 | 1 [avatar split clause 66,37,49]
68. ~v1_xboole_0(sK1) <- (~1) [resolution 31,38] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
69. ~3 | 1 [avatar split clause 68,37,45]
71. ~v1_xboole_0(sK0) <- (~1) [resolution 32,38] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
73. $false <- (~1, 2) [subsumption resolution 71,43] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
74. 1 | ~2 [avatar contradiction clause 73]
79. v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1) [resolution 33,39] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
81. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~4) [subsumption resolution 79,50] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
82. v1_xboole_0(sK0) <- (1, ~3, ~4) [subsumption resolution 81,46] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
83. $false <- (1, ~2, ~3, ~4) [subsumption resolution 82,42] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
84. ~1 | 2 | 3 | 4 [avatar contradiction clause 83]
85. $false [avatar sat refutation 52,53,54,55,67,69,74,84] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:54,thDist:-54}
% SZS output end Proof for gate_1__t14_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 10
% Generated clauses: 31
% Active clauses: 12
% Passive clauses: 12
% Final active clauses: 10
% Main loop iterations started: 12
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% 
% >>> Generating Inferences
% Binary resolution: 7
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 10
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 2
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

