i: [1,172,0,2,1,1,0,0] t14_int_2
a: [3,290,0,2,1,172]
s: 290
i: [1,174,0,7,2,1,0,0] t14_int_2
a: [3,272,0,3,1,174]
s: 272
i: [1,173,0,7,2,1,0,0] t14_int_2
a: [3,281,0,3,1,173]
s: 281
r: 272
r: 290
r: 281
s: 290
s: 272
a: [3,280,0,3,1,173]
s: 280
a: [3,285,0,4,1,173]
s: 285
i: [1,175,0,2,1,0,0,1] spc1_boole
a: [3,295,0,2,1,175]
s: 295
i: [1,176,0,2,1,0,0,3] fc1_xboole_0
a: [3,300,0,2,1,176]
s: 300
i: [1,185,0,2,1,0,0,4] fc6_ordinal1
a: [3,335,0,2,1,185]
s: 335
i: [1,184,0,2,1,0,0,4] fc6_ordinal1
a: [3,340,0,2,1,184]
s: 340
i: [1,186,0,2,1,0,0,255] spc1_numerals
a: [3,360,0,2,1,186]
s: 360
i: [1,240,0,7,3,0,0,1] reflexivity_r1_int_1
a: [3,364,0,2,1,240]
s: 364
d: [2,522,1,0,0,2,40,364,290]
e: 522
d: [2,521,1,0,0,2,52,290,364]
e: 521
r: 364
i: [1,241,0,2,1,0,0,2] rc2_xboole_0
a: [3,372,0,2,1,241]
s: 372
i: [1,242,0,2,1,0,0,1] rc2_int_1
a: [3,377,0,2,1,242]
s: 377
i: [1,243,0,2,1,0,0,2] rc4_ordinal1
a: [3,382,0,2,1,243]
s: 382
i: [1,244,0,2,1,0,0,3] rc1_ordinal1
a: [3,387,0,2,1,244]
s: 387
i: [1,245,0,2,1,0,0,2] rc1_xreal_0
a: [3,392,0,2,1,245]
s: 392
i: [1,246,0,2,1,0,0,2] rc1_xboole_0
a: [3,397,0,2,1,246]
s: 397
i: [1,250,0,2,1,0,0,255] rc2_ordinal1
a: [3,402,0,2,1,250]
s: 402
i: [1,249,0,2,1,0,0,255] rc2_ordinal1
a: [3,407,0,2,1,249]
s: 407
i: [1,248,0,2,1,0,0,255] rc2_ordinal1
a: [3,412,0,2,1,248]
s: 412
i: [1,247,0,2,1,0,0,255] rc2_ordinal1
a: [3,417,0,2,1,247]
s: 417
i: [1,252,0,2,1,0,0,2] rc5_ordinal1
a: [3,422,0,2,1,252]
s: 422
i: [1,251,0,2,1,0,0,2] rc5_ordinal1
a: [3,427,0,2,1,251]
s: 427
i: [1,255,0,2,1,0,0,4] rc3_xreal_0
a: [3,432,0,2,1,255]
s: 432
i: [1,254,0,2,1,0,0,4] rc3_xreal_0
a: [3,437,0,2,1,254]
s: 437
i: [1,253,0,2,1,0,0,4] rc3_xreal_0
a: [3,442,0,2,1,253]
s: 442
i: [1,258,0,2,1,0,0,255] rc2_xreal_0
a: [3,447,0,2,1,258]
s: 447
i: [1,257,0,2,1,0,0,255] rc2_xreal_0
a: [3,452,0,2,1,257]
s: 452
i: [1,256,0,2,1,0,0,255] rc2_xreal_0
a: [3,457,0,2,1,256]
s: 457
i: [1,261,0,2,1,0,0,5] rc4_xreal_0
a: [3,462,0,2,1,261]
s: 462
i: [1,260,0,2,1,0,0,5] rc4_xreal_0
a: [3,467,0,2,1,260]
s: 467
i: [1,259,0,2,1,0,0,5] rc4_xreal_0
a: [3,472,0,2,1,259]
s: 472
i: [1,265,0,2,1,0,0,5] rc1_int_1
a: [3,477,0,2,1,265]
s: 477
i: [1,264,0,2,1,0,0,5] rc1_int_1
a: [3,482,0,2,1,264]
s: 482
i: [1,263,0,2,1,0,0,5] rc1_int_1
a: [3,487,0,2,1,263]
s: 487
i: [1,268,0,2,1,0,0,3] rc2_nat_1
a: [3,497,0,2,1,268]
s: 497
i: [1,267,0,2,1,0,0,3] rc2_nat_1
a: [3,502,0,2,1,267]
s: 502
i: [1,234,0,3,1,0,0,255] reflexivity_r1_tarski
s: 234
i: [1,177,0,3,1,0,0,2] redefinition_k6_numbers
a: [3,305,0,3,1,177]
s: 305
i: [1,178,0,3,1,0,0,4] redefinition_k5_numbers
a: [3,310,0,3,1,178]
s: 310
i: [1,179,0,3,1,0,0,3] t1_numerals
a: [3,315,0,3,1,179]
s: 315
i: [1,189,0,3,1,0,0,255] spc1_numerals
a: [3,345,0,3,1,189]
s: 345
i: [1,188,0,3,1,0,0,255] spc1_numerals
a: [3,350,0,3,1,188]
s: 350
i: [1,262,0,3,1,0,0,5] rc1_int_1
a: [3,492,0,3,1,262]
s: 492
d: [2,525,0,3,1,2,22,350,310]
a: [3,529,0,3,1,525]
s: 529
i: [1,191,0,4,2,0,0,1] cc3_int_1
s: 191
d: [2,531,1,2,1,1,40,191,290]
a: [3,537,1,2,1,531]
s: 537
i: [1,197,0,4,2,0,0,1] cc2_int_1
s: 197
i: [1,198,0,4,2,0,0,2] cc6_ordinal1
s: 198
d: [2,532,1,2,1,1,40,191,377]
a: [3,542,1,2,1,532]
s: 542
i: [1,199,0,4,2,0,0,2] cc2_xreal_0
s: 199
d: [2,544,1,2,1,1,40,197,382]
a: [3,549,1,2,1,544]
s: 549
i: [1,201,0,4,2,0,0,4] cc3_nat_1
s: 201
i: [1,202,0,4,2,0,0,2] cc1_nat_1
s: 202
d: [2,545,1,2,1,1,40,197,422]
a: [3,554,1,2,1,545]
s: 554
i: [1,204,0,4,2,0,0,255] cc1_ordinal1
s: 204
d: [2,556,1,2,1,1,40,198,382]
a: [3,561,1,2,1,556]
s: 561
i: [1,205,0,4,2,0,0,255] cc1_ordinal1
s: 205
i: [1,207,0,4,2,0,0,2] cc3_xreal_0
s: 207
d: [2,640,2,2,1,1,40,207,537]
a: [3,650,2,2,1,640]
s: 650
i: [1,212,0,4,2,0,0,255] cc9_ordinal1
s: 212
d: [2,557,1,2,1,1,40,198,422]
a: [3,566,1,2,1,557]
s: 566
i: [1,213,0,4,2,0,0,2] cc7_ordinal1
s: 213
i: [1,214,0,4,2,0,0,3] cc3_ordinal1
s: 214
d: [2,568,1,2,1,1,40,199,382]
a: [3,573,1,2,1,568]
s: 573
i: [1,233,0,4,1,0,0,4] existence_m1_subset_1
s: 233
d: [2,569,1,2,1,1,40,199,422]
a: [3,578,1,2,1,569]
s: 578
i: [1,180,0,4,1,0,0,3] dt_k5_numbers
a: [3,320,0,4,1,180]
s: 320
i: [1,181,0,4,1,0,0,2] dt_k6_numbers
a: [3,325,0,4,1,181]
s: 325
d: [2,582,1,2,1,1,40,201,437]
a: [3,586,1,2,1,582]
s: 586
i: [1,187,0,4,1,0,0,255] spc1_numerals
a: [3,355,0,4,1,187]
s: 355
d: [2,590,1,2,1,1,40,204,335]
a: [3,597,1,2,1,590]
s: 597
i: [1,266,0,4,1,0,0,3] rc2_nat_1
a: [3,507,0,4,1,266]
s: 507
d: [2,740,0,4,1,2,22,320,310]
a: [3,744,0,4,1,740]
s: 744
d: [2,591,1,2,1,1,40,204,387]
a: [3,602,1,2,1,591]
s: 602
d: [2,746,0,4,1,2,22,325,310]
d: [2,747,0,4,1,3,22,746,305]
a: [3,751,0,4,1,747]
s: 751
d: [2,593,1,2,1,1,40,204,497]
a: [3,607,1,2,1,593]
s: 607
d: [2,753,0,4,1,2,22,355,310]
a: [3,757,0,4,1,753]
s: 757
i: [1,192,0,5,2,0,0,1] fc3_int_1
s: 192
d: [2,615,1,2,1,1,40,205,335]
a: [3,623,1,2,1,615]
s: 623
i: [1,193,0,5,2,0,0,1] fc3_int_1
s: 193
d: [2,617,1,2,1,1,40,205,387]
a: [3,633,1,2,1,617]
s: 633
i: [1,194,0,5,2,0,0,1] l8_int_2
s: 194
i: [1,195,0,5,2,0,0,1] l8_int_2
s: 195
d: [2,619,1,2,1,1,40,205,497]
a: [3,638,1,2,1,619]
s: 638
i: [1,208,0,5,2,0,0,1] fc3_xreal_0
s: 208
d: [2,642,1,2,1,1,40,207,392]
a: [3,660,1,2,1,642]
s: 660
i: [1,209,0,5,2,0,0,1] fc3_xreal_0
s: 209
i: [1,210,0,5,2,0,0,1] dt_k4_xcmplx_0
s: 210
d: [2,662,1,2,1,1,40,212,300]
a: [3,668,1,2,1,662]
s: 668
i: [1,215,0,5,2,0,0,3] t6_boole
s: 215
d: [2,663,1,2,1,1,40,212,397]
a: [3,673,1,2,1,663]
s: 673
i: [1,216,0,5,2,0,0,4] cc8_ordinal1
s: 216
i: [1,217,0,5,2,0,0,3] cc2_nat_1
s: 217
d: [2,664,1,2,1,1,40,212,472]
a: [3,678,1,2,1,664]
s: 678
i: [1,218,0,5,2,0,0,3] cc1_xreal_0
s: 218
d: [2,692,1,2,1,1,40,213,300]
a: [3,698,1,2,1,692]
s: 698
i: [1,230,0,5,2,0,0,255] t7_boole
s: 230
i: [1,182,0,5,1,0,0,1] rqRealNeg__k4_xcmplx_0__rm1_r1
a: [3,330,0,5,1,182]
s: 330
r: 280
r: 285
d: [2,510,0,2,1,2,23,290,281]
a: [3,519,0,2,1,510]
s: 519
s: 281
r: 650
r: 537
r: 272
r: 290
r: 281
s: 290
s: 272
s: 280
s: 537
d: [2,509,0,3,1,2,23,272,281]
a: [3,514,0,3,1,509]
s: 514
d: [2,838,1,0,0,2,52,519,514,195]
e: 838
d: [2,839,1,2,1,1,40,514,195]
d: [2,841,1,0,0,2,21,839,519]
e: 841
r: 519
a: [3,518,0,2,1,510]
s: 518
s: 285
d: [2,815,0,4,1,2,22,330,285]
a: [3,819,0,4,1,815]
s: 819
d: [2,854,1,4,2,1,44,193,819]
d: [2,861,1,2,1,2,21,854,518]
d: [2,862,1,0,0,3,21,861,290]
e: 862
r: 514
r: 518
s: 519
a: [3,513,0,3,1,509]
s: 513
s: 650
s: 330
a: [3,367,0,5,2,240]
s: 367
d: [2,823,2,2,1,2,23,650,281]
a: [3,830,2,2,1,823]
s: 830
d: [2,762,0,5,2,1,22,194,305]
s: 762
d: [2,693,1,2,1,1,40,213,397]
a: [3,703,1,2,1,693]
s: 703
d: [2,786,0,5,2,1,22,217,310]
s: 786
i: [1,196,0,6,2,0,0,1] l8_int_2
s: 196
d: [2,905,0,5,2,1,22,196,285]
s: 905
d: [2,906,1,0,0,3,52,519,272,905]
e: 906
d: [2,907,1,2,1,2,40,905,272]
d: [2,909,1,0,0,3,21,907,519]
e: 909
r: 272
a: [3,276,0,4,1,174]
d: [2,911,0,3,1,2,22,276,285]
a: [3,915,0,3,1,911]
s: 915
d: [2,917,1,0,0,3,52,290,915,905]
e: 917
d: [2,918,1,0,0,3,52,290,915,367]
e: 918
d: [2,919,1,2,1,2,40,915,905]
d: [2,923,1,0,0,3,21,919,290]
e: 923
d: [2,920,1,2,1,2,40,915,367]
d: [2,925,1,0,0,3,21,920,290]
e: 925
r: 280
r: 905
r: 285
r: 819
r: 915
a: [3,271,0,3,1,174]
s: 271
s: 281
r: 271
r: 290
r: 537
r: 650
a: [3,914,0,3,1,911]
s: 914
s: 276
a: [3,284,0,4,1,173]
s: 284
a: [3,818,0,4,1,815]
s: 818
d: [2,934,0,4,1,2,22,276,281]
a: [3,938,0,4,1,934]
s: 938
d: [2,940,0,4,1,2,22,284,281]
a: [3,944,0,4,1,940]
s: 944
d: [2,948,1,3,1,1,40,938,195]
a: [3,952,1,3,1,948]
s: 952
d: [2,954,2,0,0,2,52,519,952,193]
e: 954
d: [2,955,2,2,1,1,40,952,193]
d: [2,957,2,0,0,2,21,955,519]
e: 957
d: [2,825,0,3,1,2,23,272,281]
d: [2,865,1,2,1,2,21,854,290]
f: 957,948,934,925,911,909,865,839,825,815,521,510,240,182,172,173,174
% Refutation found. Thanks to Tanya!
% SZS status Theorem for int_2__t14_int_2
% SZS output start Proof for int_2__t14_int_2
1. ! [X0] : (v1_int_1(X0) => ((k4_xcmplx_0(1) = X0 | 1 = X0) => (r1_int_1(X0,k4_xcmplx_0(1)) & r1_int_1(X0,1)))) [input t14_int_2]
2. ~! [X0] : (v1_int_1(X0) => ((k4_xcmplx_0(1) = X0 | 1 = X0) => (r1_int_1(X0,k4_xcmplx_0(1)) & r1_int_1(X0,1)))) [negated conjecture 1]
62. ! [X0,X1] : ((v1_int_1(X1) & v1_int_1(X0)) => r1_int_1(X0,X0)) [input reflexivity_r1_int_1]
67. ! [X0] : (v1_int_1(X0) => (v1_int_1(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0)))) [input fc3_int_1]
69. 1 = k4_xcmplx_0(k4_xcmplx_0(1)) [input rqRealNeg__k4_xcmplx_0__rm1_r1]
72. ! [X0] : (v1_int_1(X0) => (r1_int_1(k4_xcmplx_0(1),X0) & r1_int_1(1,X0) & r1_int_1(X0,k6_numbers))) [input l8_int_2]
89. ? [X0] : (((~r1_int_1(X0,k4_xcmplx_0(1)) | ~r1_int_1(X0,1)) & (k4_xcmplx_0(1) = X0 | 1 = X0)) & v1_int_1(X0)) [ennf transformation 2]
90. ? [X0] : ((~r1_int_1(X0,k4_xcmplx_0(1)) | ~r1_int_1(X0,1)) & (k4_xcmplx_0(1) = X0 | 1 = X0) & v1_int_1(X0)) [flattening 89]
93. ! [X0] : ((v1_int_1(k4_xcmplx_0(X0)) & v1_xcmplx_0(k4_xcmplx_0(X0))) | ~v1_int_1(X0)) [ennf transformation 67]
94. ! [X0] : ((r1_int_1(k4_xcmplx_0(1),X0) & r1_int_1(1,X0) & r1_int_1(X0,k6_numbers)) | ~v1_int_1(X0)) [ennf transformation 72]
137. ! [X0,X1] : (r1_int_1(X0,X0) | (~v1_int_1(X1) | ~v1_int_1(X0))) [ennf transformation 62]
138. ! [X0,X1] : (r1_int_1(X0,X0) | ~v1_int_1(X1) | ~v1_int_1(X0)) [flattening 137]
139. ? [X0] : ((~r1_int_1(X0,k4_xcmplx_0(1)) | ~r1_int_1(X0,1)) & (k4_xcmplx_0(1) = X0 | 1 = X0) & v1_int_1(X0)) => ((~r1_int_1(sK0,k4_xcmplx_0(1)) | ~r1_int_1(sK0,1)) & (k4_xcmplx_0(1) = sK0 | 1 = sK0) & v1_int_1(sK0)) [choice axiom]
140. (~r1_int_1(sK0,k4_xcmplx_0(1)) | ~r1_int_1(sK0,1)) & (k4_xcmplx_0(1) = sK0 | 1 = sK0) & v1_int_1(sK0) [skolemisation 90,139]
172. v1_int_1(sK0) [cnf transformation 140] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
173. k4_xcmplx_0(1) = sK0 | 1 = sK0 [cnf transformation 140] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
174. ~r1_int_1(sK0,k4_xcmplx_0(1)) | ~r1_int_1(sK0,1) [cnf transformation 140] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
182. 1 = k4_xcmplx_0(k4_xcmplx_0(1)) [cnf transformation 69] {a:0,w:5,wCS:2000,sine:1,thAx:0,allAx:1,thDist:-1}
193. v1_int_1(k4_xcmplx_0(X0)) | ~v1_int_1(X0) [cnf transformation 93] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
195. r1_int_1(1,X0) | ~v1_int_1(X0) [cnf transformation 94] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
196. r1_int_1(k4_xcmplx_0(1),X0) | ~v1_int_1(X0) [cnf transformation 94] {a:0,w:6,wCS:2400,sine:1,thAx:0,allAx:1,thDist:-1}
240. r1_int_1(X0,X0) | ~v1_int_1(X1) | ~v1_int_1(X0) [cnf transformation 138] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
270. 1 <=> r1_int_1(sK0,1) [avatar definition]
272. ~r1_int_1(sK0,1) <- (~1) [avatar component clause 270] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
274. 2 <=> r1_int_1(sK0,k4_xcmplx_0(1)) [avatar definition]
276. ~r1_int_1(sK0,k4_xcmplx_0(1)) <- (~2) [avatar component clause 274] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
277. ~1 | ~2 [avatar split clause 174,274,270]
279. 3 <=> 1 = sK0 [avatar definition]
281. 1 = sK0 <- (3) [avatar component clause 279] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
283. 4 <=> k4_xcmplx_0(1) = sK0 [avatar definition]
285. k4_xcmplx_0(1) = sK0 <- (4) [avatar component clause 283] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
286. 3 | 4 [avatar split clause 173,283,279]
288. 5 <=> v1_int_1(sK0) [avatar definition]
290. v1_int_1(sK0) <- (5) [avatar component clause 288] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
291. 5 [avatar split clause 172,288]
328. 13 <=> 1 = k4_xcmplx_0(k4_xcmplx_0(1)) [avatar definition]
330. 1 = k4_xcmplx_0(k4_xcmplx_0(1)) <- (13) [avatar component clause 328] {a:0,w:5,wCS:2000,sine:255,thAx:0,allAx:1,thDist:-1}
331. 13 [avatar split clause 182,328]
363. 20 <=> ! [X1] : ~v1_int_1(X1) [avatar definition]
364. ~v1_int_1(X1) <- (20) [avatar component clause 363] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
366. 21 <=> ! [X0] : (r1_int_1(X0,X0) | ~v1_int_1(X0)) [avatar definition]
367. r1_int_1(X0,X0) | ~v1_int_1(X0) <- (21) [avatar component clause 366] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
368. 20 | 21 [avatar split clause 240,366,363]
510. v1_int_1(1) <- (3, 5) [backward demodulation 290,281] {a:0,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
512. 50 <=> r1_int_1(1,1) [avatar definition]
514. ~r1_int_1(1,1) <- (~50) [avatar component clause 512] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
517. 51 <=> v1_int_1(1) [avatar definition]
519. v1_int_1(1) <- (51) [avatar component clause 517] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
520. 51 | ~3 | ~5 [avatar split clause 510,288,279,517]
521. $false <- (5, 20) [unit resulting resolution 290,364] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
524. ~5 | ~20 [avatar contradiction clause 521]
815. 1 = k4_xcmplx_0(sK0) <- (4, 13) [forward demodulation 330,285] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
817. 95 <=> 1 = k4_xcmplx_0(sK0) [avatar definition]
819. 1 = k4_xcmplx_0(sK0) <- (95) [avatar component clause 817] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
820. 95 | ~4 | ~13 [avatar split clause 815,328,283,817]
825. ~r1_int_1(1,1) <- (~1, 3) [backward demodulation 272,281] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
833. ~50 | 1 | ~3 [avatar split clause 825,279,270,512]
839. ~v1_int_1(1) <- (~50) [resolution 514,195] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
843. ~51 | 50 [avatar split clause 839,512,517]
854. v1_int_1(1) | ~v1_int_1(sK0) <- (95) [superposition 193,819] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
865. v1_int_1(1) <- (5, 95) [subsumption resolution 854,290] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
866. 51 | ~5 | ~95 [avatar split clause 865,817,288,517]
905. r1_int_1(sK0,X0) | ~v1_int_1(X0) <- (4) [forward demodulation 196,285] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:1,thDist:-1}
907. ~v1_int_1(1) <- (~1, 4) [resolution 905,272] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
909. $false <- (~1, 4, 51) [subsumption resolution 907,519] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
910. 1 | ~4 | ~51 [avatar contradiction clause 909]
911. ~r1_int_1(sK0,sK0) <- (~2, 4) [forward demodulation 276,285] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
913. 102 <=> r1_int_1(sK0,sK0) [avatar definition]
915. ~r1_int_1(sK0,sK0) <- (~102) [avatar component clause 913] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
916. ~102 | 2 | ~4 [avatar split clause 911,283,274,913]
920. ~v1_int_1(sK0) <- (21, ~102) [resolution 915,367] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
925. $false <- (5, 21, ~102) [subsumption resolution 920,290] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
926. ~5 | ~21 | 102 [avatar contradiction clause 925]
934. ~r1_int_1(1,k4_xcmplx_0(1)) <- (~2, 3) [forward demodulation 276,281] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
936. 103 <=> r1_int_1(1,k4_xcmplx_0(1)) [avatar definition]
938. ~r1_int_1(1,k4_xcmplx_0(1)) <- (~103) [avatar component clause 936] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
939. ~103 | 2 | ~3 [avatar split clause 934,279,274,936]
948. ~v1_int_1(k4_xcmplx_0(1)) <- (~103) [resolution 938,195] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
950. 105 <=> v1_int_1(k4_xcmplx_0(1)) [avatar definition]
952. ~v1_int_1(k4_xcmplx_0(1)) <- (~105) [avatar component clause 950] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
953. ~105 | 103 [avatar split clause 948,936,950]
955. ~v1_int_1(1) <- (~105) [resolution 952,193] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
957. $false <- (51, ~105) [subsumption resolution 955,519] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
958. ~51 | 105 [avatar contradiction clause 957]
959. $false [avatar sat refutation 277,286,291,331,368,520,524,820,833,843,866,910,916,926,939,953,958] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:573,thDist:-573}
% SZS output end Proof for int_2__t14_int_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 72
% >>> Preprocessing
% Introduced skolems: 16
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 97
% Generated clauses: 363
% Active clauses: 118
% Passive clauses: 183
% Final active clauses: 94
% Final passive clauses: 50
% Main loop iterations started: 132
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 11
% Fw demodulations: 15
% Bw demodulations: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 26
% 
% >>> Generating Inferences
% Binary resolution: 70
% Unit resulting resolution: 6
% Forward superposition: 7
% Backward superposition: 14
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 14
% Sat splitting refutations: 13
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 140
% SAT solver unit clauses: 45
% SAT solver binary clauses: 60
% 
% Memory used [KB]: 5500
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s ( own 0.0000 s ) 
% term sharing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

