i: [1,258,0,2,1,1,0,0] t3_nat_1
a: [3,404,0,2,1,258]
s: 404
i: [1,260,0,3,1,1,0,0] t3_nat_1
a: [3,394,0,3,1,260]
s: 394
i: [1,259,0,3,1,1,0,0] t3_nat_1
a: [3,399,0,3,1,259]
s: 399
i: [1,261,0,2,1,0,0,255] fc8_numbers
a: [3,409,0,2,1,261]
s: 409
i: [1,262,0,2,1,0,0,3] fc7_card_1
a: [3,414,0,2,1,262]
s: 414
i: [1,263,0,2,1,0,0,255] fc1_numbers
a: [3,419,0,2,1,263]
s: 419
i: [1,264,0,2,1,0,0,255] fc4_card_1
a: [3,424,0,2,1,264]
s: 424
i: [1,265,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,429,0,2,1,265]
s: 429
i: [1,273,0,2,1,0,0,3] fc6_ordinal1
a: [3,464,0,2,1,273]
s: 464
i: [1,272,0,2,1,0,0,3] fc6_ordinal1
a: [3,469,0,2,1,272]
s: 469
i: [1,355,0,2,1,0,0,4] rc3_card_1
a: [3,481,0,2,1,355]
s: 481
i: [1,356,0,2,1,0,0,255] rc1_card_1
a: [3,486,0,2,1,356]
s: 486
i: [1,357,0,2,1,0,0,1] rc4_ordinal1
a: [3,491,0,2,1,357]
s: 491
i: [1,358,0,2,1,0,0,2] rc1_xreal_0
a: [3,496,0,2,1,358]
s: 496
i: [1,359,0,2,1,0,0,4] rc1_ordinal1
a: [3,501,0,2,1,359]
s: 501
i: [1,360,0,2,1,0,0,2] rc1_xxreal_0
a: [3,506,0,2,1,360]
s: 506
i: [1,364,0,2,1,0,0,255] rc2_ordinal1
a: [3,511,0,2,1,364]
s: 511
i: [1,363,0,2,1,0,0,255] rc2_ordinal1
a: [3,516,0,2,1,363]
s: 516
i: [1,362,0,2,1,0,0,255] rc2_ordinal1
a: [3,521,0,2,1,362]
s: 521
i: [1,361,0,2,1,0,0,255] rc2_ordinal1
a: [3,526,0,2,1,361]
s: 526
i: [1,366,0,2,1,0,0,1] rc5_ordinal1
a: [3,531,0,2,1,366]
s: 531
i: [1,365,0,2,1,0,0,1] rc5_ordinal1
a: [3,536,0,2,1,365]
s: 536
i: [1,371,0,2,1,0,0,255] rc2_card_1
a: [3,541,0,2,1,371]
s: 541
i: [1,370,0,2,1,0,0,255] rc2_card_1
a: [3,546,0,2,1,370]
s: 546
i: [1,369,0,2,1,0,0,255] rc2_card_1
a: [3,551,0,2,1,369]
s: 551
i: [1,368,0,2,1,0,0,255] rc2_card_1
a: [3,556,0,2,1,368]
s: 556
i: [1,367,0,2,1,0,0,255] rc2_card_1
a: [3,561,0,2,1,367]
s: 561
i: [1,374,0,2,1,0,0,255] rc3_xreal_0
a: [3,566,0,2,1,374]
s: 566
i: [1,373,0,2,1,0,0,255] rc3_xreal_0
a: [3,571,0,2,1,373]
s: 571
i: [1,372,0,2,1,0,0,255] rc3_xreal_0
a: [3,576,0,2,1,372]
s: 576
i: [1,377,0,2,1,0,0,255] rc2_xreal_0
a: [3,581,0,2,1,377]
s: 581
i: [1,376,0,2,1,0,0,255] rc2_xreal_0
a: [3,586,0,2,1,376]
s: 586
i: [1,375,0,2,1,0,0,255] rc2_xreal_0
a: [3,591,0,2,1,375]
s: 591
i: [1,379,0,2,1,0,0,2] rc3_xxreal_0
a: [3,596,0,2,1,379]
s: 596
i: [1,378,0,2,1,0,0,2] rc3_xxreal_0
a: [3,601,0,2,1,378]
s: 601
i: [1,381,0,2,1,0,0,2] rc2_xxreal_0
a: [3,606,0,2,1,381]
s: 606
i: [1,380,0,2,1,0,0,2] rc2_xxreal_0
a: [3,611,0,2,1,380]
s: 611
i: [1,384,0,2,1,0,0,255] rc4_xreal_0
a: [3,616,0,2,1,384]
s: 616
i: [1,383,0,2,1,0,0,255] rc4_xreal_0
a: [3,621,0,2,1,383]
s: 621
i: [1,382,0,2,1,0,0,255] rc4_xreal_0
a: [3,626,0,2,1,382]
s: 626
i: [1,386,0,2,1,0,0,2] rc4_xxreal_0
a: [3,631,0,2,1,386]
s: 631
i: [1,385,0,2,1,0,0,2] rc4_xxreal_0
a: [3,636,0,2,1,385]
s: 636
i: [1,390,0,2,1,0,0,255] rc1_real_1
a: [3,641,0,2,1,390]
s: 641
i: [1,389,0,2,1,0,0,255] rc1_real_1
a: [3,646,0,2,1,389]
s: 646
i: [1,388,0,2,1,0,0,255] rc1_real_1
a: [3,651,0,2,1,388]
s: 651
i: [1,274,0,3,1,0,0,5] fc1_subset_1
s: 274
i: [1,344,0,3,1,0,0,5] rc2_subset_1
s: 344
i: [1,347,0,3,1,0,0,255] reflexivity_r1_tarski
s: 347
i: [1,266,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,434,0,3,1,266]
s: 434
i: [1,267,0,3,1,0,0,255] d2_xboole_0
a: [3,439,0,3,1,267]
s: 439
i: [1,268,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,444,0,3,1,268]
s: 444
i: [1,269,0,3,1,0,0,2] t1_numerals
a: [3,449,0,3,1,269]
s: 449
i: [1,387,0,3,1,0,0,255] rc1_real_1
a: [3,656,0,3,1,387]
s: 656
d: [2,658,0,3,1,2,22,439,434]
a: [3,662,0,3,1,658]
s: 662
r: 429
d: [2,670,0,2,1,2,23,429,662]
a: [3,674,0,2,1,670]
s: 674
d: [2,664,0,3,1,2,22,449,434]
a: [3,668,0,3,1,664]
s: 668
i: [1,285,0,4,2,0,0,255] cc1_card_1
s: 285
i: [1,287,0,4,2,0,0,255] cc3_card_1
s: 287
d: [2,684,1,2,1,1,40,287,404]
a: [3,690,1,2,1,684]
s: 690
i: [1,288,0,4,2,0,0,4] cc5_card_1
s: 288
i: [1,289,0,4,2,0,0,1] cc2_xreal_0
s: 289
d: [2,708,1,2,1,1,40,288,404]
a: [3,714,1,2,1,708]
s: 714
i: [1,290,0,4,2,0,0,4] cc6_ordinal1
s: 290
d: [2,726,1,2,1,1,40,289,404]
a: [3,732,1,2,1,726]
s: 732
i: [1,291,0,4,2,0,0,1] cc2_xxreal_0
s: 291
i: [1,293,0,4,2,0,0,2] cc4_xreal_0
s: 293
d: [2,758,1,2,1,1,40,291,404]
a: [3,764,1,2,1,758]
s: 764
i: [1,302,0,4,2,0,0,255] cc1_ordinal1
s: 302
d: [2,702,2,2,1,1,40,690,285]
a: [3,706,2,2,1,702]
s: 706
i: [1,303,0,4,2,0,0,255] cc1_ordinal1
s: 303
i: [1,305,0,4,2,0,0,255] cc9_ordinal1
s: 305
d: [2,802,3,2,1,1,40,706,302]
a: [3,806,3,2,1,802]
s: 806
i: [1,306,0,4,2,0,0,255] cc2_card_1
s: 306
d: [2,809,3,2,1,1,40,303,706]
a: [3,821,3,2,1,809]
s: 821
i: [1,307,0,4,2,0,0,1] cc7_ordinal1
s: 307
i: [1,308,0,4,2,0,0,4] cc3_ordinal1
s: 308
d: [2,677,1,2,1,1,40,285,486]
a: [3,682,1,2,1,677]
s: 682
i: [1,342,0,4,1,0,0,3] existence_m1_subset_1
s: 342
d: [2,685,1,2,1,1,40,287,491]
a: [3,695,1,2,1,685]
s: 695
i: [1,346,0,4,1,0,0,255] rc3_subset_1
s: 346
i: [1,270,0,4,1,0,0,255] dt_k5_numbers
a: [3,454,0,4,1,270]
s: 454
d: [2,686,1,2,1,1,40,287,531]
a: [3,700,1,2,1,686]
s: 700
i: [1,271,0,4,1,0,0,255] dt_k6_numbers
a: [3,459,0,4,1,271]
s: 459
d: [2,709,1,2,1,1,40,288,491]
a: [3,719,1,2,1,709]
s: 719
d: [2,918,0,4,1,2,22,454,444]
a: [3,922,0,4,1,918]
s: 922
d: [2,926,0,4,1,2,22,459,444]
a: [3,930,0,4,1,926]
s: 930
d: [2,710,1,2,1,1,40,288,531]
a: [3,724,1,2,1,710]
s: 724
i: [1,275,0,5,2,0,0,4] fc9_card_1
s: 275
d: [2,727,1,2,1,1,40,289,491]
a: [3,737,1,2,1,727]
s: 737
i: [1,277,0,5,2,0,0,4] rc4_card_1
s: 277
i: [1,281,0,5,2,0,0,5] rc1_subset_1
s: 281
d: [2,728,1,2,1,1,40,289,531]
a: [3,742,1,2,1,728]
s: 742
i: [1,292,0,5,2,0,0,1] t2_nat_1
s: 292
d: [2,745,1,2,1,1,40,290,491]
a: [3,751,1,2,1,745]
s: 751
i: [1,309,0,5,2,0,0,1] t6_boole
s: 309
i: [1,310,0,5,2,0,0,255] cc8_card_1
s: 310
d: [2,746,1,2,1,1,40,290,531]
a: [3,756,1,2,1,746]
s: 756
i: [1,313,0,5,2,0,0,255] cc7_card_1
s: 313
d: [2,759,1,2,1,1,40,291,491]
a: [3,769,1,2,1,759]
s: 769
i: [1,314,0,5,2,0,0,255] cc1_xreal_0
s: 314
i: [1,315,0,5,2,0,0,255] cc1_real_1
s: 315
d: [2,760,1,2,1,1,40,291,531]
a: [3,774,1,2,1,760]
s: 774
i: [1,316,0,5,2,0,0,3] cc4_card_1
s: 316
d: [2,777,1,2,1,1,40,293,496]
a: [3,786,1,2,1,777]
s: 786
i: [1,317,0,5,2,0,0,3] cc8_ordinal1
s: 317
i: [1,339,0,5,2,0,0,255] t7_boole
s: 339
d: [2,788,1,2,1,1,40,302,464]
a: [3,795,1,2,1,788]
s: 795
i: [1,343,0,5,1,0,0,5] rc2_subset_1
s: 343
d: [2,789,1,2,1,1,40,302,501]
a: [3,800,1,2,1,789]
s: 800
i: [1,345,0,5,1,0,0,255] rc3_subset_1
s: 345
i: [1,353,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,476,0,5,2,353]
s: 476
d: [2,808,1,2,1,1,40,303,464]
a: [3,816,1,2,1,808]
s: 816
d: [2,948,0,5,2,1,22,309,434]
s: 948
d: [2,995,1,5,2,2,44,394,948]
s: 995
d: [2,949,0,5,2,1,22,310,434]
s: 949
d: [2,962,0,5,2,1,22,313,434]
s: 962
d: [2,810,1,2,1,1,40,303,501]
a: [3,826,1,2,1,810]
s: 826
i: [1,283,0,6,2,0,0,255] rc4_subset_1
s: 283
d: [2,828,1,2,1,1,40,305,674]
a: [3,835,1,2,1,828]
s: 835
i: [1,286,0,6,2,0,0,255] rc5_card_1
s: 286
i: [1,322,0,6,3,0,0,255] cc2_ordinal1
s: 322
d: [2,830,1,2,1,1,40,305,626]
a: [3,840,1,2,1,830]
s: 840
i: [1,323,0,6,3,0,0,4] cc6_card_1
s: 323
d: [2,831,1,2,1,1,40,305,636]
a: [3,845,1,2,1,831]
s: 845
i: [1,326,0,6,3,0,0,2] cc5_xxreal_0
s: 326
i: [1,328,0,6,3,0,0,2] cc5_xxreal_0
s: 328
d: [2,847,1,2,1,1,40,306,674]
a: [3,854,1,2,1,847]
s: 854
i: [1,329,0,6,3,0,0,2] cc3_xxreal_0
s: 329
d: [2,849,1,2,1,1,40,306,626]
a: [3,859,1,2,1,849]
s: 859
i: [1,331,0,6,3,0,0,2] cc3_xxreal_0
s: 331
i: [1,333,0,6,3,0,0,2] cc7_xxreal_0
s: 333
d: [2,850,1,2,1,1,40,306,636]
a: [3,864,1,2,1,850]
s: 864
i: [1,334,0,6,3,0,0,2] cc7_xxreal_0
s: 334
d: [2,866,1,2,1,1,40,307,674]
a: [3,873,1,2,1,866]
s: 873
i: [1,335,0,6,2,0,0,255] t1_subset
s: 335
i: [1,348,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 348
d: [2,868,1,2,1,1,40,307,626]
a: [3,878,1,2,1,868]
s: 878
i: [1,276,0,7,2,0,0,4] rc4_card_1
s: 276
d: [2,869,1,2,1,1,40,307,636]
a: [3,883,1,2,1,869]
s: 883
i: [1,280,0,7,2,0,0,5] rc1_subset_1
s: 280
i: [1,282,0,7,2,0,0,255] rc4_subset_1
s: 282
d: [2,885,1,2,1,1,40,308,674]
a: [3,892,1,2,1,885]
s: 892
i: [1,304,0,7,3,0,0,4] cc5_ordinal1
s: 304
d: [2,887,1,2,1,1,40,308,626]
a: [3,897,1,2,1,887]
s: 897
i: [1,337,0,7,2,0,0,255] t3_subset
s: 337
i: [1,338,0,7,3,0,0,1] t8_boole
s: 338
d: [2,888,1,2,1,1,40,308,636]
a: [3,902,1,2,1,888]
s: 902
i: [1,284,0,8,3,0,0,255] cc10_ordinal1
s: 284
d: [2,970,1,2,1,1,40,316,668]
a: [3,975,1,2,1,970]
s: 975
i: [1,311,0,8,3,0,0,5] cc1_subset_1
s: 311
i: [1,319,0,8,4,0,0,2] cc4_xxreal_0
s: 319
d: [2,1205,2,6,3,1,40,319,764]
a: [3,1220,2,2,1,1205]
s: 1220
d: [2,1279,3,0,0,3,52,674,399,1220,338]
e: 1279
d: [2,1280,3,0,0,3,52,674,399,1220,338]
e: 1280
d: [2,1281,3,0,0,3,52,399,1220,948]
e: 1281
r: 1220
i: [1,321,0,8,4,0,0,2] cc6_xxreal_0
s: 321
a: [3,1219,2,2,1,1205]
s: 1219
i: [1,324,0,8,4,0,0,2] cc8_xxreal_0
s: 324
i: [1,336,0,8,3,0,0,255] t2_subset
s: 336
a: [3,1224,2,2,1,1205]
s: 1224
i: [1,312,0,9,3,0,0,255] cc4_subset_1
s: 312
d: [2,1014,1,4,2,1,40,323,546]
d: [2,1015,1,2,1,2,21,1014,551]
a: [3,1019,1,2,1,1015]
s: 1019
i: [1,341,0,9,3,0,0,255] t5_subset
s: 341
i: [1,340,0,10,3,0,0,255] t4_subset
s: 340
d: [2,1021,1,4,2,1,40,326,571]
d: [2,1023,1,2,1,2,21,1021,576]
a: [3,1027,1,2,1,1023]
s: 1027
i: [1,354,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 354
d: [2,1022,1,4,2,1,40,326,596]
d: [2,1029,1,2,1,2,21,1022,601]
a: [3,1033,1,2,1,1029]
s: 1033
i: [1,278,0,11,4,0,0,255] cc3_subset_1
s: 278
i: [1,279,0,11,4,0,0,255] cc2_subset_1
s: 279
d: [2,1035,1,4,2,1,40,328,571]
d: [2,1037,1,2,1,2,21,1035,576]
a: [3,1041,1,2,1,1037]
s: 1041
i: [1,296,0,11,5,0,0,1] t2_real
s: 296
d: [2,1377,1,8,4,1,40,296,394]
d: [2,1383,1,6,3,2,21,1377,732]
a: [3,1386,1,2,1,1383]
s: 1386
r: 1386
i: [1,297,0,11,5,0,0,1] t1_real
s: 297
r: 1224
i: [1,298,0,11,5,0,0,1] t8_real
s: 298
a: [3,1387,1,2,1,1383]
s: 1387
i: [1,299,0,11,5,0,0,1] t7_real
s: 299
a: [3,1223,2,2,1,1205]
s: 1223
i: [1,300,0,11,5,0,0,1] t3_real
s: 300
i: [1,301,0,11,5,0,0,1] t4_real
s: 301
a: [3,1228,2,2,1,1205]
s: 1228
i: [1,294,0,13,6,0,0,1] t5_real
s: 294
d: [2,1401,2,4,2,1,40,1386,328]
a: [3,1407,2,2,1,1401]
s: 1407
i: [1,295,0,13,6,0,0,1] t6_real
s: 295
i: [1,352,0,14,4,0,0,255] existence_m2_subset_1
s: 352
d: [2,1440,2,4,2,2,44,1387,948]
s: 1440
r: 326
d: [2,1474,0,12,3,0,21,352,311]
s: 1474
d: [2,1464,3,4,2,2,44,1407,948]
s: 1464
r: 329
i: [1,349,0,15,5,0,0,255] dt_m2_subset_1
s: 349
d: [2,1484,0,13,4,0,21,349,311]
s: 1484
d: [2,1445,1,12,6,0,40,294,292]
d: [2,1451,1,10,5,1,21,1445,1387]
d: [2,1452,1,8,4,1,21,1451,289]
a: [3,1455,1,6,3,1452]
s: 1455
d: [2,1495,3,0,0,4,52,1223,1219,404,1455]
e: 1495
d: [2,1497,2,4,2,2,40,1455,404]
d: [2,1504,2,2,1,3,21,1497,1219]
d: [2,1505,2,0,0,4,21,1504,1223]
e: 1505
r: 1223
r: 1464
r: 1407
r: 1228
s: 329
s: 1224
i: [1,350,0,15,5,0,0,255] redefinition_m2_subset_1
s: 350
d: [2,1538,0,13,4,0,21,350,311]
s: 1538
a: [3,1227,2,2,1,1205]
s: 1227
i: [1,351,0,15,5,0,0,255] redefinition_m2_subset_1
s: 351
a: [3,1406,2,2,1,1401]
s: 1406
d: [2,982,1,2,1,1,40,317,668]
d: [2,1077,2,2,1,1,40,873,291]
d: [2,1079,2,2,1,1,40,873,289]
d: [2,1402,2,4,2,1,40,1386,326]
d: [2,1409,2,2,1,2,21,1402,674]
d: [2,1411,1,8,4,1,40,297,394]
d: [2,1434,1,6,3,2,21,1411,732]
d: [2,1543,3,4,2,1,40,1406,329]
d: [2,1545,3,2,1,2,21,1543,674]
f: 1545,1504,1452,1434,1409,1281,1079,1077,982,726,670,664,658,269,267,266,265,258,259,260
% Refutation found. Thanks to Tanya!
% SZS status Theorem for nat_1__t3_nat_1
% SZS output start Proof for nat_1__t3_nat_1
1. ! [X0] : (v7_ordinal1(X0) => ~(r1_xxreal_0(X0,k6_numbers) & k6_numbers != X0)) [input t3_nat_1]
2. ~! [X0] : (v7_ordinal1(X0) => ~(r1_xxreal_0(X0,k6_numbers) & k6_numbers != X0)) [negated conjecture 1]
7. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
9. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
21. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
27. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v2_xxreal_0(X1) | v3_xxreal_0(X0) | v1_xboole_0(X1))))) [input t5_real]
36. v1_xboole_0(o_0_0_xboole_0) [input dt_o_0_0_xboole_0]
51. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
67. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
91. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
92. k1_xboole_0 = o_0_0_xboole_0 [input d2_xboole_0]
95. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
97. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
98. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
103. ! [X0] : (v7_ordinal1(X0) => r1_xxreal_0(k6_numbers,X0)) [input t2_nat_1]
120. ? [X0] : ((r1_xxreal_0(X0,k6_numbers) & k6_numbers != X0) & v7_ordinal1(X0)) [ennf transformation 2]
121. ? [X0] : (r1_xxreal_0(X0,k6_numbers) & k6_numbers != X0 & v7_ordinal1(X0)) [flattening 120]
135. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 97]
137. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 98]
138. ! [X0] : (r1_xxreal_0(k6_numbers,X0) | ~v7_ordinal1(X0)) [ennf transformation 103]
140. ! [X0] : (! [X1] : (((v2_xxreal_0(X1) | v3_xxreal_0(X0) | v1_xboole_0(X1)) | ~r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 27]
141. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | v3_xxreal_0(X0) | v1_xboole_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 140]
146. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 21]
147. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 146]
162. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 91]
170. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 51]
181. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 9]
182. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 181]
183. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 7]
184. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 183]
207. ? [X0] : (r1_xxreal_0(X0,k6_numbers) & k6_numbers != X0 & v7_ordinal1(X0)) => (r1_xxreal_0(sK0,k6_numbers) & k6_numbers != sK0 & v7_ordinal1(sK0)) [choice axiom]
208. r1_xxreal_0(sK0,k6_numbers) & k6_numbers != sK0 & v7_ordinal1(sK0) [skolemisation 121,207]
258. v7_ordinal1(sK0) [cnf transformation 208] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
259. k6_numbers != sK0 [cnf transformation 208] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
260. r1_xxreal_0(sK0,k6_numbers) [cnf transformation 208] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
265. v1_xboole_0(o_0_0_xboole_0) [cnf transformation 36] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
266. k6_numbers = k1_xboole_0 [cnf transformation 95] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
267. k1_xboole_0 = o_0_0_xboole_0 [cnf transformation 92] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
269. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 67] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
289. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 135] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
291. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 137] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
292. r1_xxreal_0(k6_numbers,X0) | ~v7_ordinal1(X0) [cnf transformation 138] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
294. ~r1_xxreal_0(X0,X1) | v3_xxreal_0(X0) | v1_xboole_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 141] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
297. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 147] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
309. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 162] {a:0,w:5,wCS:2000,sine:1,thAx:0,allAx:1,thDist:-1}
317. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 170] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
326. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 182] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
329. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 184] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
392. 1 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
394. r1_xxreal_0(sK0,k6_numbers) <- (1) [avatar component clause 392] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
395. 1 [avatar split clause 260,392]
397. 2 <=> k6_numbers = sK0 [avatar definition]
399. k6_numbers != sK0 <- (~2) [avatar component clause 397] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
400. ~2 [avatar split clause 259,397]
402. 3 <=> v7_ordinal1(sK0) [avatar definition]
404. v7_ordinal1(sK0) <- (3) [avatar component clause 402] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
405. 3 [avatar split clause 258,402]
427. 8 <=> v1_xboole_0(o_0_0_xboole_0) [avatar definition]
429. v1_xboole_0(o_0_0_xboole_0) <- (8) [avatar component clause 427] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
430. 8 [avatar split clause 265,427]
432. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
434. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 432] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
435. 9 [avatar split clause 266,432]
437. 10 <=> k1_xboole_0 = o_0_0_xboole_0 [avatar definition]
439. k1_xboole_0 = o_0_0_xboole_0 <- (10) [avatar component clause 437] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
440. 10 [avatar split clause 267,437]
447. 12 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
449. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (12) [avatar component clause 447] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
450. 12 [avatar split clause 269,447]
658. k6_numbers = o_0_0_xboole_0 <- (9, 10) [forward demodulation 439,434] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
660. 55 <=> k6_numbers = o_0_0_xboole_0 [avatar definition]
662. k6_numbers = o_0_0_xboole_0 <- (55) [avatar component clause 660] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
663. 55 | ~9 | ~10 [avatar split clause 658,437,432,660]
664. m1_subset_1(k6_numbers,k4_ordinal1) <- (9, 12) [forward demodulation 449,434] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
666. 56 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
668. m1_subset_1(k6_numbers,k4_ordinal1) <- (56) [avatar component clause 666] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
669. 56 | ~9 | ~12 [avatar split clause 664,447,432,666]
670. v1_xboole_0(k6_numbers) <- (8, 55) [backward demodulation 429,662] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
672. 57 <=> v1_xboole_0(k6_numbers) [avatar definition]
674. v1_xboole_0(k6_numbers) <- (57) [avatar component clause 672] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
675. 57 | ~8 | ~55 [avatar split clause 670,660,427,672]
726. v1_xreal_0(sK0) <- (3) [resolution 289,404] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
730. 66 <=> v1_xreal_0(sK0) [avatar definition]
732. v1_xreal_0(sK0) <- (66) [avatar component clause 730] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
733. 66 | ~3 [avatar split clause 726,402,730]
871. 87 <=> v7_ordinal1(k6_numbers) [avatar definition]
873. v7_ordinal1(k6_numbers) <- (87) [avatar component clause 871] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
948. k6_numbers = X0 | ~v1_xboole_0(X0) <- (9) [forward demodulation 309,434] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
982. v7_ordinal1(k6_numbers) <- (56) [resolution 317,668] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
984. 87 | ~56 [avatar split clause 982,666,871]
1077. v1_xxreal_0(k6_numbers) <- (87) [resolution 873,291] {a:2,w:2,wCS:800,sine:1,thAx:0,allAx:3,thDist:-3}
1079. v1_xreal_0(k6_numbers) <- (87) [resolution 873,289] {a:2,w:2,wCS:800,sine:1,thAx:0,allAx:3,thDist:-3}
1084. 114 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1087. 114 | ~87 [avatar split clause 1077,871,1084]
1090. 115 <=> v1_xreal_0(k6_numbers) [avatar definition]
1093. 115 | ~87 [avatar split clause 1079,871,1090]
1218. 129 <=> v1_xboole_0(sK0) [avatar definition]
1219. ~v1_xboole_0(sK0) <- (~129) [avatar component clause 1218] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1220. v1_xboole_0(sK0) <- (129) [avatar component clause 1218] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1222. 130 <=> v2_xxreal_0(sK0) [avatar definition]
1281. $false <- (~2, 9, 129) [unit resulting resolution 399,1220,948] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1288. 2 | ~9 | ~129 [avatar contradiction clause 1281]
1385. 148 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1386. v3_xxreal_0(k6_numbers) <- (148) [avatar component clause 1385] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1387. ~v3_xxreal_0(k6_numbers) <- (~148) [avatar component clause 1385] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1402. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (148) [resolution 1386,326] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1405. 151 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1406. v2_xxreal_0(k6_numbers) <- (151) [avatar component clause 1405] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1409. ~v1_xxreal_0(k6_numbers) <- (57, 148) [subsumption resolution 1402,674] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
1410. ~114 | ~57 | ~148 [avatar split clause 1409,1385,672,1084]
1411. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (1) [resolution 297,394] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1434. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (1, 66) [subsumption resolution 1411,732] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1435. ~115 | 151 | ~130 | ~1 | ~66 [avatar split clause 1434,730,392,1222,1405,1090]
1445. v3_xxreal_0(k6_numbers) | v1_xboole_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(k6_numbers) | ~v7_ordinal1(X1) [resolution 294,292] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
1451. v1_xboole_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(k6_numbers) | ~v7_ordinal1(X1) <- (~148) [subsumption resolution 1445,1387] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1452. v1_xboole_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(k6_numbers) | ~v7_ordinal1(X1) <- (~148) [subsumption resolution 1451,289] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1454. 154 <=> ! [X1] : (v1_xboole_0(X1) | ~v7_ordinal1(X1) | v2_xxreal_0(X1)) [avatar definition]
1455. ~v7_ordinal1(X1) | v1_xboole_0(X1) | v2_xxreal_0(X1) <- (154) [avatar component clause 1454] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1456. ~115 | 154 | 148 [avatar split clause 1452,1385,1454,1090]
1497. v1_xboole_0(sK0) | v2_xxreal_0(sK0) <- (3, 154) [resolution 1455,404] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1504. v2_xxreal_0(sK0) <- (3, ~129, 154) [subsumption resolution 1497,1219] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1508. 130 | ~3 | 129 | ~154 [avatar split clause 1504,1454,1218,402,1222]
1543. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (151) [resolution 1406,329] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1545. ~v1_xxreal_0(k6_numbers) <- (57, 151) [subsumption resolution 1543,674] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1546. ~114 | ~57 | ~151 [avatar split clause 1545,1405,672,1084]
1547. $false [avatar sat refutation 395,400,405,430,435,440,450,663,669,675,733,984,1087,1093,1288,1410,1435,1456,1508,1546] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:906,thDist:-906}
% SZS output end Proof for nat_1__t3_nat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 102
% >>> Preprocessing
% Introduced skolems: 25
% Pure predicates: 3
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 133
% Generated clauses: 645
% Active clauses: 173
% Passive clauses: 302
% Final active clauses: 163
% Final passive clauses: 106
% Main loop iterations started: 194
% 
% >>> Simplifying Inferences
% Duplicate literals: 19
% Fw subsumption resolutions: 55
% Fw demodulations: 7
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 19
% Forward subsumptions: 85
% Backward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 250
% Unit resulting resolution: 4
% Factoring: 1
% Forward superposition: 11
% Backward superposition: 5
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 24
% Split components: 57
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 183
% SAT solver unit clauses: 55
% SAT solver binary clauses: 73
% 
% Memory used [KB]: 5756
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% passive container maintenance: 0.002 s
% other: 0.007 s ( own 0.001 s ) 

