i: [1,55,0,2,1,1,0,0] t111_funct_4
a: [3,86,0,2,1,55]
s: 86
i: [1,54,0,2,1,1,0,0] t111_funct_4
a: [3,91,0,2,1,54]
s: 91
i: [1,56,0,7,1,1,0,0] t111_funct_4
a: [3,81,0,7,1,56]
s: 81
i: [1,77,0,2,1,0,0,1] rc1_funct_1
a: [3,104,0,2,1,77]
s: 104
i: [1,76,0,2,1,0,0,1] rc1_funct_1
a: [3,109,0,2,1,76]
s: 109
i: [1,64,0,3,1,0,0,255] reflexivity_r1_tarski
s: 64
i: [1,63,0,4,1,0,0,255] existence_m1_subset_1
s: 63
i: [1,65,0,6,2,0,0,1] dt_k5_relat_1
s: 65
i: [1,60,0,7,2,0,0,255] t59_relat_1
s: 60
i: [1,61,0,7,2,0,0,255] t3_subset
s: 61
i: [1,62,0,7,2,0,0,255] t3_subset
s: 62
d: [2,111,1,5,1,0,40,61,63]
s: 111
i: [1,57,0,8,3,0,0,255] cc2_relat_1
s: 57
d: [2,113,1,6,2,0,40,57,63]
s: 113
i: [1,74,0,8,3,0,0,1] fc8_funct_1
s: 74
i: [1,75,0,8,3,0,0,1] fc8_funct_1
s: 75
d: [2,114,1,7,3,0,40,57,62]
s: 114
i: [1,67,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,98,0,9,3,67]
s: 98
i: [1,58,0,10,4,0,0,255] cc3_funct_1
s: 58
i: [1,66,0,11,2,0,0,1] rd5_relat_1
s: 66
i: [1,70,0,12,5,0,0,1] dt_k1_funct_4
s: 70
d: [2,119,1,8,3,0,40,58,63]
s: 119
i: [1,71,0,12,5,0,0,1] dt_k1_funct_4
s: 71
d: [2,120,1,9,4,0,40,58,62]
s: 120
i: [1,72,0,12,5,0,0,1] dt_k1_funct_4
s: 72
i: [1,73,0,12,5,0,0,1] dt_k1_funct_4
s: 73
i: [1,59,0,16,6,0,0,255] t97_funct_4
s: 59
d: [2,133,0,14,5,0,21,59,120]
d: [2,134,0,12,4,0,21,133,114]
s: 134
d: [2,135,1,12,4,1,44,81,134]
d: [2,140,1,9,3,1,19,135]
d: [2,141,1,7,2,2,21,140,86]
d: [2,142,1,5,1,3,21,141,91]
a: [3,146,1,5,1,142]
s: 146
d: [2,148,2,0,0,2,52,91,146,60]
e: 148
d: [2,149,2,2,1,1,40,146,60]
d: [2,151,2,0,0,2,21,149,91]
e: 151
f: 151,142,54,55,56
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t111_funct_4
% SZS output start Proof for funct_4__t111_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : k1_funct_4(k5_relat_1(X0,X1),X0) = X0) [input t111_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : k1_funct_4(k5_relat_1(X0,X1),X0) = X0) [negated conjecture 1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
7. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_funct_1(X1))) [input cc3_funct_1]
18. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
19. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(X0,X1) => k1_funct_4(X0,X1) = X1))) [input t97_funct_4]
20. ! [X0] : ! [X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [input t59_relat_1]
21. ! [X0,X1] : (v1_relat_1(X1) => r1_tarski(k5_relat_1(X1,X0),X1)) [flattening 20]
22. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 18]
24. ? [X0] : (? [X1] : k1_funct_4(k5_relat_1(X0,X1),X0) != X0 & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
25. ? [X0] : (? [X1] : k1_funct_4(k5_relat_1(X0,X1),X0) != X0 & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 24]
26. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 6]
27. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 7]
28. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 27]
29. ! [X0] : (! [X1] : ((k1_funct_4(X0,X1) = X1 | ~r1_tarski(X0,X1)) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 19]
30. ! [X0] : (! [X1] : (k1_funct_4(X0,X1) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 29]
31. ! [X0,X1] : (r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1)) [ennf transformation 21]
46. ? [X0] : (? [X1] : k1_funct_4(k5_relat_1(X0,X1),X0) != X0 & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : sK0 != k1_funct_4(k5_relat_1(sK0,X1),sK0) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
47. ? [X1] : sK0 != k1_funct_4(k5_relat_1(sK0,X1),sK0) => sK0 != k1_funct_4(k5_relat_1(sK0,sK1),sK0) [choice axiom]
48. sK0 != k1_funct_4(k5_relat_1(sK0,sK1),sK0) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 25,47,46]
49. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 22]
54. v1_relat_1(sK0) [cnf transformation 48] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
55. v1_funct_1(sK0) [cnf transformation 48] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
56. sK0 != k1_funct_4(k5_relat_1(sK0,sK1),sK0) [cnf transformation 48] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
57. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
58. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_funct_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
59. k1_funct_4(X0,X1) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:16,wCS:6400,sine:255,thAx:0,allAx:1,thDist:-1}
60. r1_tarski(k5_relat_1(X1,X0),X1) | ~v1_relat_1(X1) [cnf transformation 31] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
62. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 49] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
79. 1 <=> sK0 = k1_funct_4(k5_relat_1(sK0,sK1),sK0) [avatar definition]
81. sK0 != k1_funct_4(k5_relat_1(sK0,sK1),sK0) <- (~1) [avatar component clause 79] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. ~1 [avatar split clause 56,79]
84. 2 <=> v1_funct_1(sK0) [avatar definition]
86. v1_funct_1(sK0) <- (2) [avatar component clause 84] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
87. 2 [avatar split clause 55,84]
89. 3 <=> v1_relat_1(sK0) [avatar definition]
91. v1_relat_1(sK0) <- (3) [avatar component clause 89] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. 3 [avatar split clause 54,89]
114. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 57,62] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
120. ~r1_tarski(X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | v1_funct_1(X1) [resolution 58,62] {a:1,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
133. k1_funct_4(X0,X1) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [subsumption resolution 59,120] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
134. k1_funct_4(X0,X1) = X1 | ~r1_tarski(X0,X1) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [subsumption resolution 133,114] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
135. sK0 != sK0 | ~r1_tarski(k5_relat_1(sK0,sK1),sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [superposition 81,134] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
140. ~r1_tarski(k5_relat_1(sK0,sK1),sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [trivial inequality removal 135] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
141. ~r1_tarski(k5_relat_1(sK0,sK1),sK0) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 140,86] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
142. ~r1_tarski(k5_relat_1(sK0,sK1),sK0) <- (~1, 2, 3) [subsumption resolution 141,91] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
144. 8 <=> r1_tarski(k5_relat_1(sK0,sK1),sK0) [avatar definition]
146. ~r1_tarski(k5_relat_1(sK0,sK1),sK0) <- (~8) [avatar component clause 144] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
147. ~8 | 1 | ~2 | ~3 [avatar split clause 142,89,84,79,144]
149. ~v1_relat_1(sK0) <- (~8) [resolution 146,60] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
151. $false <- (3, ~8) [subsumption resolution 149,91] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
152. ~3 | 8 [avatar contradiction clause 151]
153. $false [avatar sat refutation 82,87,92,147,152] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:45,thDist:-45}
% SZS output end Proof for funct_4__t111_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 19
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 66
% Active clauses: 25
% Passive clauses: 30
% Final active clauses: 25
% Final passive clauses: 1
% Main loop iterations started: 29
% 
% >>> Simplifying Inferences
% Duplicate literals: 14
% Trivial inequalities: 1
% Fw subsumption resolutions: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 13
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 6
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 5
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

