i: [1,34,0,2,1,1,0,0] t182_relat_1
a: [3,59,0,2,1,34]
s: 59
i: [1,36,0,3,1,1,0,0] t182_relat_1
a: [3,49,0,3,1,36]
s: 49
i: [1,35,0,3,1,1,0,0] t182_relat_1
a: [3,54,0,3,1,35]
s: 54
i: [1,33,0,3,1,1,0,0] t182_relat_1
a: [3,64,0,3,1,33]
s: 64
i: [1,44,0,3,1,0,0,1] reflexivity_r1_tarski
s: 44
i: [1,43,0,4,1,0,0,2] existence_m1_subset_1
s: 43
i: [1,40,0,7,2,0,0,2] t3_subset
s: 40
i: [1,41,0,7,2,0,0,2] t3_subset
s: 41
d: [2,66,1,5,1,0,40,40,43]
s: 66
i: [1,37,0,8,3,0,0,2] cc2_relat_1
s: 37
i: [1,38,0,9,3,0,0,255] d18_relat_1
s: 38
d: [2,68,1,6,2,0,40,37,43]
s: 68
i: [1,39,0,9,3,0,0,255] d18_relat_1
s: 39
d: [2,71,1,6,2,0,40,39,44]
s: 71
i: [1,42,0,9,3,0,0,1] t1_xboole_1
s: 42
i: [1,45,0,12,4,0,0,1] cc5_relat_1
s: 45
d: [2,75,1,6,2,1,40,42,64]
s: 75
d: [2,69,1,7,3,0,40,37,41]
s: 69
d: [2,89,2,4,2,1,40,69,64]
a: [3,98,2,2,1,89]
s: 98
d: [2,77,1,10,3,0,40,45,43]
s: 77
d: [2,78,1,11,4,0,40,45,41]
s: 78
d: [2,81,2,5,1,1,40,75,66]
a: [3,85,2,5,1,81]
s: 85
d: [2,74,1,12,4,0,40,42,38]
s: 74
d: [2,100,2,8,3,1,40,78,54]
d: [2,104,2,6,2,2,21,100,59]
s: 104
d: [2,80,2,9,3,1,40,75,38]
s: 80
d: [2,76,2,8,2,0,40,42,66]
s: 76
d: [2,113,3,5,1,2,40,104,66]
a: [3,117,3,5,1,113]
s: 117
d: [2,101,2,11,4,0,40,78,71]
d: [2,103,2,9,3,0,18,101]
s: 103
d: [2,126,4,12,3,1,40,117,78]
a: [3,130,4,4,1,126]
s: 130
d: [2,137,5,0,0,2,52,59,130,68]
e: 137
d: [2,138,5,0,0,2,52,59,66,130,69]
e: 138
d: [2,139,5,0,0,2,52,59,43,130,37]
e: 139
d: [2,140,5,2,1,1,40,130,68]
d: [2,144,5,0,0,2,21,140,59]
e: 144
r: 130
d: [2,88,2,10,4,0,40,69,38]
s: 88
a: [3,93,2,2,1,89]
s: 93
a: [3,129,4,4,1,126]
s: 129
d: [2,110,2,11,3,0,40,74,66]
s: 110
d: [2,107,3,8,2,1,40,85,42]
s: 107
d: [2,102,2,17,5,0,40,78,77]
d: [2,105,2,13,4,0,21,102,68]
s: 105
d: [2,109,2,15,5,0,40,74,38]
s: 109
d: [2,173,4,7,1,1,40,107,66]
a: [3,178,4,7,1,173]
s: 178
d: [2,119,3,10,4,1,40,80,39]
d: [2,122,3,8,3,1,18,119]
s: 122
d: [2,193,4,0,0,4,52,59,49,54,122]
e: 193
d: [2,194,4,5,2,2,40,122,54]
d: [2,198,4,3,1,3,21,194,59]
d: [2,199,4,0,0,4,21,198,49]
e: 199
f: 199,33,34,35,36
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t182_relat_1
% SZS output start Proof for relat_1__t182_relat_1
1. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => ! [X2] : ((v4_relat_1(X2,X0) & v1_relat_1(X2)) => v4_relat_1(X2,X1))) [input t182_relat_1]
2. ~! [X0] : ! [X1] : (r1_tarski(X0,X1) => ! [X2] : ((v4_relat_1(X2,X0) & v1_relat_1(X2)) => v4_relat_1(X2,X1))) [negated conjecture 1]
5. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0))) [input d18_relat_1]
11. ! [X0] : ! [X1] : ! [X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [input t1_xboole_1]
13. ~! [X0,X1] : (r1_tarski(X0,X1) => ! [X2] : ((v4_relat_1(X2,X0) & v1_relat_1(X2)) => v4_relat_1(X2,X1))) [flattening 2]
14. ! [X0,X1] : (v1_relat_1(X1) => (v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0))) [flattening 5]
16. ! [X0,X1,X2] : ((r1_tarski(X1,X2) & r1_tarski(X0,X1)) => r1_tarski(X0,X2)) [flattening 11]
18. ? [X0,X1] : (? [X2] : (~v4_relat_1(X2,X1) & (v4_relat_1(X2,X0) & v1_relat_1(X2))) & r1_tarski(X0,X1)) [ennf transformation 13]
19. ? [X0,X1] : (? [X2] : (~v4_relat_1(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & r1_tarski(X0,X1)) [flattening 18]
21. ! [X0,X1] : ((v4_relat_1(X1,X0) <=> r1_tarski(k9_xtuple_0(X1),X0)) | ~v1_relat_1(X1)) [ennf transformation 14]
22. ! [X0,X1,X2] : (r1_tarski(X0,X2) | (~r1_tarski(X1,X2) | ~r1_tarski(X0,X1))) [ennf transformation 16]
23. ! [X0,X1,X2] : (r1_tarski(X0,X2) | ~r1_tarski(X1,X2) | ~r1_tarski(X0,X1)) [flattening 22]
26. ? [X0,X1] : (? [X2] : (~v4_relat_1(X2,X1) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & r1_tarski(X0,X1)) => (? [X2] : (~v4_relat_1(X2,sK1) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & r1_tarski(sK0,sK1)) [choice axiom]
27. ? [X2] : (~v4_relat_1(X2,sK1) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (~v4_relat_1(sK2,sK1) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
28. (~v4_relat_1(sK2,sK1) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & r1_tarski(sK0,sK1) [skolemisation 19,27,26]
29. ! [X0,X1] : (((v4_relat_1(X1,X0) | ~r1_tarski(k9_xtuple_0(X1),X0)) & (r1_tarski(k9_xtuple_0(X1),X0) | ~v4_relat_1(X1,X0))) | ~v1_relat_1(X1)) [nnf transformation 21]
33. r1_tarski(sK0,sK1) [cnf transformation 28] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
34. v1_relat_1(sK2) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
35. v4_relat_1(sK2,sK0) [cnf transformation 28] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
36. ~v4_relat_1(sK2,sK1) [cnf transformation 28] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. r1_tarski(k9_xtuple_0(X1),X0) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 29] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. ~r1_tarski(k9_xtuple_0(X1),X0) | v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 29] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. ~r1_tarski(X1,X2) | r1_tarski(X0,X2) | ~r1_tarski(X0,X1) [cnf transformation 23] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. 1 <=> v4_relat_1(sK2,sK1) [avatar definition]
49. ~v4_relat_1(sK2,sK1) <- (~1) [avatar component clause 47] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
50. ~1 [avatar split clause 36,47]
52. 2 <=> v4_relat_1(sK2,sK0) [avatar definition]
54. v4_relat_1(sK2,sK0) <- (2) [avatar component clause 52] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
55. 2 [avatar split clause 35,52]
57. 3 <=> v1_relat_1(sK2) [avatar definition]
59. v1_relat_1(sK2) <- (3) [avatar component clause 57] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
60. 3 [avatar split clause 34,57]
62. 4 <=> r1_tarski(sK0,sK1) [avatar definition]
64. r1_tarski(sK0,sK1) <- (4) [avatar component clause 62] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
65. 4 [avatar split clause 33,62]
75. ~r1_tarski(X5,sK0) | r1_tarski(X5,sK1) <- (4) [resolution 42,64] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
80. r1_tarski(k9_xtuple_0(X0),sK1) | ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) <- (4) [resolution 75,38] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
119. ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | v4_relat_1(X0,sK1) | ~v1_relat_1(X0) <- (4) [resolution 80,39] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
122. ~v4_relat_1(X0,sK0) | ~v1_relat_1(X0) | v4_relat_1(X0,sK1) <- (4) [duplicate literal removal 119] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
194. ~v1_relat_1(sK2) | v4_relat_1(sK2,sK1) <- (2, 4) [resolution 122,54] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
198. v4_relat_1(sK2,sK1) <- (2, 3, 4) [subsumption resolution 194,59] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
199. $false <- (~1, 2, 3, 4) [subsumption resolution 198,49] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
200. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 199]
203. $false [avatar sat refutation 50,55,60,65,200] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:182,thDist:-182}
% SZS output end Proof for relat_1__t182_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 114
% Active clauses: 38
% Passive clauses: 60
% Final active clauses: 37
% Final passive clauses: 22
% Main loop iterations started: 38
% 
% >>> Simplifying Inferences
% Duplicate literals: 7
% Fw subsumption resolutions: 13
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 66
% Unit resulting resolution: 4
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 9
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 19
% SAT solver unit clauses: 4
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

