i: [1,134,0,2,1,1,0,0] t35_interva1
a: [3,204,0,2,1,134]
s: 204
i: [1,132,0,2,1,1,0,0] t35_interva1
a: [3,214,0,2,1,132]
s: 214
i: [1,131,0,2,1,1,0,0] t35_interva1
a: [3,219,0,2,1,131]
s: 219
i: [1,135,0,3,1,1,0,0] t35_interva1
a: [3,199,0,3,1,135]
s: 199
i: [1,133,0,3,1,1,0,0] t35_interva1
a: [3,209,0,3,1,133]
s: 209
i: [1,137,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,224,0,2,1,137]
s: 224
i: [1,138,0,2,1,0,0,255] fc1_xboole_0
a: [3,229,0,2,1,138]
s: 229
i: [1,187,0,2,1,0,0,1] rc2_xboole_0
a: [3,239,0,2,1,187]
s: 239
i: [1,188,0,2,1,0,0,1] rc1_xboole_0
a: [3,244,0,2,1,188]
s: 244
i: [1,136,0,10,1,1,0,0] t35_interva1
a: [3,194,0,10,1,136]
s: 194
i: [1,140,0,3,1,0,0,2] fc1_subset_1
s: 140
i: [1,169,0,3,1,0,0,2] rc2_subset_1
s: 169
i: [1,173,0,3,1,0,0,255] rc2_interva1
s: 173
i: [1,175,0,3,1,0,0,255] reflexivity_r1_tarski
s: 175
i: [1,139,0,3,1,0,0,255] d2_xboole_0
a: [3,234,0,3,1,139]
s: 234
r: 224
i: [1,166,0,4,1,0,0,1] existence_m1_interva1
s: 166
i: [1,167,0,4,1,0,0,2] existence_m1_subset_1
s: 167
i: [1,171,0,4,1,0,0,255] rc3_subset_1
s: 171
i: [1,174,0,4,1,0,0,255] rc2_interva1
s: 174
i: [1,149,0,5,2,0,0,1] rc1_interva1
s: 149
i: [1,151,0,5,2,0,0,2] rc1_subset_1
s: 151
i: [1,154,0,5,2,0,0,255] t6_boole
s: 154
i: [1,163,0,5,2,0,0,255] t7_boole
s: 163
i: [1,168,0,5,1,0,0,2] rc2_subset_1
s: 168
i: [1,170,0,5,1,0,0,255] rc3_subset_1
s: 170
i: [1,148,0,6,2,0,0,1] rc1_interva1
s: 148
i: [1,153,0,6,2,0,0,255] rc4_subset_1
s: 153
i: [1,157,0,6,2,0,0,255] t1_subset
s: 157
i: [1,172,0,6,1,0,0,255] rc2_interva1
s: 172
i: [1,178,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 178
i: [1,150,0,7,2,0,0,2] rc1_subset_1
s: 150
i: [1,152,0,7,2,0,0,255] rc4_subset_1
s: 152
i: [1,161,0,7,2,0,0,255] t3_subset
s: 161
i: [1,162,0,7,3,0,0,2] t8_boole
s: 162
i: [1,176,0,7,1,0,0,255] commutativity_k2_setfam_1
s: 176
i: [1,177,0,7,1,0,0,255] commutativity_k3_setfam_1
s: 177
d: [2,247,1,7,3,0,44,154,154]
s: 247
i: [1,155,0,8,3,0,0,2] cc1_subset_1
s: 155
d: [2,254,1,5,2,0,40,155,170]
s: 254
i: [1,158,0,8,2,0,0,1] dt_m1_interva1
s: 158
i: [1,159,0,8,3,0,0,255] t2_subset
s: 159
d: [2,262,1,8,2,0,40,159,168]
d: [2,269,1,5,1,0,21,262,140]
s: 269
i: [1,156,0,9,3,0,0,255] cc4_subset_1
s: 156
d: [2,263,1,8,2,0,40,159,170]
d: [2,270,1,5,1,0,21,263,140]
s: 270
i: [1,165,0,9,3,0,0,255] t5_subset
s: 165
i: [1,144,0,10,4,0,0,255] l40_interva1
s: 144
d: [2,294,1,7,3,1,40,144,199]
d: [2,299,1,5,2,2,21,294,204]
d: [2,300,1,3,1,3,21,299,219]
a: [3,304,1,3,1,300]
s: 304
i: [1,164,0,10,3,0,0,255] t4_subset
s: 164
d: [2,295,1,7,3,1,40,144,209]
d: [2,306,1,5,2,2,21,295,214]
d: [2,307,1,3,1,3,21,306,219]
a: [3,311,1,3,1,307]
s: 311
i: [1,141,0,11,4,0,0,255] cc3_subset_1
s: 141
i: [1,142,0,11,4,0,0,255] cc2_subset_1
s: 142
d: [2,328,1,9,3,0,40,141,170]
d: [2,334,1,5,2,0,21,328,171]
s: 334
i: [1,189,0,16,6,0,0,1] redefinition_r1_interva1
d: [2,190,0,11,4,0,18,189]
s: 190
d: [2,272,2,5,1,0,40,269,178]
s: 272
i: [1,145,0,12,4,0,0,255] l40_interva1
s: 145
i: [1,179,0,16,6,0,0,1] reflexivity_r1_interva1
s: 179
d: [2,283,2,5,1,0,40,270,178]
s: 283
i: [1,180,0,17,6,0,0,1] fc1_interva1
s: 180
d: [2,252,1,6,2,0,40,155,167]
s: 252
i: [1,181,0,17,6,0,0,1] fc2_interva1
s: 181
i: [1,182,0,18,6,0,0,1] dt_k4_interva1
s: 182
d: [2,261,1,6,2,0,40,159,167]
s: 261
i: [1,183,0,18,6,0,0,1] dt_k3_interva1
s: 183
d: [2,264,1,10,2,0,40,159,172]
d: [2,271,1,6,1,0,21,264,140]
s: 271
i: [1,185,0,19,7,0,0,1] redefinition_r1_interva1
s: 185
i: [1,146,0,20,6,0,0,255] d4_interva1
s: 146
d: [2,350,1,21,6,1,44,194,146]
d: [2,355,1,18,5,2,21,350,199]
d: [2,356,1,16,4,3,21,355,204]
d: [2,357,1,13,3,4,21,356,209]
d: [2,358,1,11,2,5,21,357,214]
d: [2,359,1,9,1,6,21,358,219]
a: [3,363,1,9,1,359]
s: 363
i: [1,147,0,20,6,0,0,255] d3_interva1
s: 147
d: [2,365,1,27,6,1,44,194,147]
d: [2,371,1,24,5,2,21,365,209]
d: [2,372,1,22,4,3,21,371,214]
d: [2,373,1,20,3,4,21,372,219]
a: [3,381,1,6,1,373]
s: 381
d: [2,403,2,0,0,6,52,219,214,204,209,199,381,182]
e: 403
d: [2,404,2,12,5,1,40,381,182]
d: [2,407,2,9,4,2,21,404,199]
d: [2,408,2,7,3,3,21,407,204]
d: [2,409,2,4,2,4,21,408,209]
d: [2,410,2,2,1,5,21,409,214]
d: [2,411,2,0,0,6,21,410,219]
e: 411
r: 381
i: [1,184,0,20,7,0,0,1] symmetry_r1_interva1
s: 184
i: [1,160,0,27,7,0,0,255] t33_interva1
s: 160
a: [3,380,1,6,1,373]
s: 380
a: [3,385,1,9,1,373]
s: 385
d: [2,366,2,24,6,1,44,363,147]
d: [2,387,2,21,5,2,21,366,209]
d: [2,388,2,19,4,3,21,387,214]
d: [2,389,2,17,3,4,21,388,219]
a: [3,396,2,5,1,389]
s: 396
d: [2,280,1,6,2,0,40,156,168]
s: 280
d: [2,291,1,6,2,0,40,165,168]
s: 291
a: [3,393,2,4,1,389]
s: 393
d: [2,292,1,6,2,0,40,165,170]
s: 292
d: [2,422,2,13,3,1,40,380,144]
d: [2,424,2,11,2,2,21,422,219]
a: [3,428,2,6,1,424]
s: 428
d: [2,297,1,11,4,0,40,144,148]
d: [2,298,1,9,3,0,18,297]
d: [2,313,1,6,2,0,21,298,149]
s: 313
d: [2,327,1,9,3,0,40,141,168]
d: [2,333,1,6,2,0,21,327,169]
s: 333
d: [2,443,3,11,3,1,40,396,144]
d: [2,444,3,9,2,2,21,443,219]
a: [3,448,3,5,1,444]
s: 448
d: [2,249,1,7,3,0,40,155,161]
s: 249
a: [3,401,2,8,1,389]
s: 401
d: [2,257,1,10,3,0,40,159,161]
d: [2,265,1,7,2,0,21,257,140]
s: 265
d: [2,259,1,10,3,0,40,159,150]
d: [2,267,1,7,2,0,21,259,140]
s: 267
d: [2,461,3,24,7,1,44,401,160]
d: [2,462,3,22,6,2,21,461,204]
d: [2,463,3,20,5,3,21,462,214]
a: [3,470,3,4,1,463]
s: 470
d: [2,478,4,0,0,4,52,219,214,209,470,190]
e: 478
d: [2,479,4,7,3,1,40,470,190]
d: [2,481,4,4,2,2,21,479,209]
d: [2,482,4,2,1,3,21,481,214]
d: [2,483,4,0,0,4,21,482,219]
e: 483
r: 470
d: [2,260,1,10,3,0,40,159,152]
d: [2,268,1,7,2,0,21,260,140]
s: 268
a: [3,469,3,4,1,463]
s: 469
d: [2,279,1,7,2,0,40,156,167]
s: 279
d: [2,290,1,7,2,0,40,165,167]
s: 290
a: [3,466,3,16,4,463]
s: 466
d: [2,319,1,7,2,0,40,164,168]
s: 319
d: [2,493,4,14,4,1,40,466,158]
s: 493
d: [2,320,1,7,2,0,40,164,170]
s: 320
d: [2,258,1,12,3,0,40,159,158]
d: [2,266,1,8,2,0,21,258,140]
s: 266
d: [2,496,5,12,4,1,40,493,158]
s: 496
d: [2,502,6,0,0,5,52,209,199,311,304,496]
e: 502
d: [2,503,6,9,3,2,40,496,304]
d: [2,505,6,6,2,3,21,503,311]
d: [2,506,6,3,1,4,21,505,199]
d: [2,507,6,0,0,5,21,506,209]
e: 507
r: 385
r: 401
r: 496
r: 493
r: 466
a: [3,377,1,5,1,373]
s: 377
d: [2,509,2,0,0,6,52,219,214,204,209,199,377,181]
e: 509
d: [2,510,2,12,5,1,40,377,181]
d: [2,513,2,9,4,2,21,510,199]
d: [2,514,2,7,3,3,21,513,204]
d: [2,515,2,4,2,4,21,514,209]
d: [2,516,2,2,1,5,21,515,214]
d: [2,517,2,0,0,6,21,516,219]
e: 517
d: [2,436,2,20,6,1,44,385,146]
d: [2,437,2,17,5,2,21,436,199]
d: [2,438,2,15,4,3,21,437,204]
d: [2,439,2,12,3,4,21,438,209]
d: [2,440,2,10,2,5,21,439,214]
d: [2,441,2,8,1,6,21,440,219]
f: 517,507,483,463,441,411,373,307,300,131,132,133,134,135,136
% Refutation found. Thanks to Tanya!
% SZS status Theorem for interva1__t35_interva1
% SZS output start Proof for interva1__t35_interva1
1. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2)) => r1_interva1(X0,k3_interva1(X0,X1,k4_interva1(X0,X1,X2)),X1)))) [input t35_interva1]
2. ~! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2)) => r1_interva1(X0,k3_interva1(X0,X1,k4_interva1(X0,X1,X2)),X1)))) [negated conjecture 1]
11. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2)) => k3_interva1(X0,X1,X2) = k3_setfam_1(X1,X2)))) [input d3_interva1]
12. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2)) => k4_interva1(X0,X1,X2) = k2_setfam_1(X1,X2)))) [input d4_interva1]
18. ! [X0,X1,X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2) & m1_interva1(X1,X0) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => m1_interva1(k4_interva1(X0,X1,X2),X0)) [input dt_k4_interva1]
19. ! [X0] : ! [X1] : (m1_interva1(X1,X0) => m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [input dt_m1_interva1]
27. ! [X0,X1,X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2) & m1_interva1(X1,X0) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k4_interva1(X0,X1,X2))) [input fc2_interva1]
28. ! [X0] : (~v1_xboole_0(X0) => ! [X1] : ((m1_interva1(X1,X0) & ~v1_xboole_0(X1)) => (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)))) [input l40_interva1]
37. ! [X0,X1,X2] : ((m1_interva1(X2,X0) & ~v1_xboole_0(X2) & m1_interva1(X1,X0) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => (r1_interva1(X0,X1,X2) <=> X1 = X2)) [input redefinition_r1_interva1]
43. ! [X0] : ! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X2,X0) & ~v1_xboole_0(X2)) => k3_setfam_1(X1,k2_setfam_1(X1,X2)) = X1)) [input t33_interva1]
51. ! [X0,X1] : (m1_interva1(X1,X0) => m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [flattening 19]
53. ! [X0,X1] : ((m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)) => ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X2,X0) & ~v1_xboole_0(X2)) => k3_setfam_1(X1,k2_setfam_1(X1,X2)) = X1)) [flattening 43]
61. ? [X0] : (? [X1] : (? [X2] : (~r1_interva1(X0,k3_interva1(X0,X1,k4_interva1(X0,X1,X2)),X1) & (m1_interva1(X2,X0) & ~v1_xboole_0(X2))) & (m1_interva1(X1,X0) & ~v1_xboole_0(X1))) & ~v1_xboole_0(X0)) [ennf transformation 2]
62. ? [X0] : (? [X1] : (? [X2] : (~r1_interva1(X0,k3_interva1(X0,X1,k4_interva1(X0,X1,X2)),X1) & m1_interva1(X2,X0) & ~v1_xboole_0(X2)) & m1_interva1(X1,X0) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) [flattening 61]
67. ! [X0] : (! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)) | (~m1_interva1(X1,X0) | v1_xboole_0(X1))) | v1_xboole_0(X0)) [ennf transformation 28]
68. ! [X0] : (! [X1] : ((m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) & v1_interva1(X1,X0) & ~v1_xboole_0(X1)) | ~m1_interva1(X1,X0) | v1_xboole_0(X1)) | v1_xboole_0(X0)) [flattening 67]
69. ! [X0] : (! [X1] : (! [X2] : (k4_interva1(X0,X1,X2) = k2_setfam_1(X1,X2) | (~m1_interva1(X2,X0) | v1_xboole_0(X2))) | (~m1_interva1(X1,X0) | v1_xboole_0(X1))) | v1_xboole_0(X0)) [ennf transformation 12]
70. ! [X0] : (! [X1] : (! [X2] : (k4_interva1(X0,X1,X2) = k2_setfam_1(X1,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2)) | ~m1_interva1(X1,X0) | v1_xboole_0(X1)) | v1_xboole_0(X0)) [flattening 69]
71. ! [X0] : (! [X1] : (! [X2] : (k3_interva1(X0,X1,X2) = k3_setfam_1(X1,X2) | (~m1_interva1(X2,X0) | v1_xboole_0(X2))) | (~m1_interva1(X1,X0) | v1_xboole_0(X1))) | v1_xboole_0(X0)) [ennf transformation 11]
72. ! [X0] : (! [X1] : (! [X2] : (k3_interva1(X0,X1,X2) = k3_setfam_1(X1,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2)) | ~m1_interva1(X1,X0) | v1_xboole_0(X1)) | v1_xboole_0(X0)) [flattening 71]
80. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~m1_interva1(X1,X0)) [ennf transformation 51]
83. ! [X0,X1] : (! [X2] : (k3_setfam_1(X1,k2_setfam_1(X1,X2)) = X1 | (~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X2,X0) | v1_xboole_0(X2))) | (~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X1,X0) | v1_xboole_0(X1))) [ennf transformation 53]
84. ! [X0,X1] : (! [X2] : (k3_setfam_1(X1,k2_setfam_1(X1,X2)) = X1 | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X2,X0) | v1_xboole_0(X2)) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X1,X0) | v1_xboole_0(X1)) [flattening 83]
96. ! [X0,X1,X2] : (~v1_xboole_0(k4_interva1(X0,X1,X2)) | (~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 27]
97. ! [X0,X1,X2] : (~v1_xboole_0(k4_interva1(X0,X1,X2)) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 96]
98. ! [X0,X1,X2] : (m1_interva1(k4_interva1(X0,X1,X2),X0) | (~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 18]
99. ! [X0,X1,X2] : (m1_interva1(k4_interva1(X0,X1,X2),X0) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 98]
104. ! [X0,X1,X2] : ((r1_interva1(X0,X1,X2) <=> X1 = X2) | (~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 37]
105. ! [X0,X1,X2] : ((r1_interva1(X0,X1,X2) <=> X1 = X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 104]
106. ? [X0] : (? [X1] : (? [X2] : (~r1_interva1(X0,k3_interva1(X0,X1,k4_interva1(X0,X1,X2)),X1) & m1_interva1(X2,X0) & ~v1_xboole_0(X2)) & m1_interva1(X1,X0) & ~v1_xboole_0(X1)) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : (~r1_interva1(sK0,k3_interva1(sK0,X1,k4_interva1(sK0,X1,X2)),X1) & m1_interva1(X2,sK0) & ~v1_xboole_0(X2)) & m1_interva1(X1,sK0) & ~v1_xboole_0(X1)) & ~v1_xboole_0(sK0)) [choice axiom]
107. ? [X1] : (? [X2] : (~r1_interva1(sK0,k3_interva1(sK0,X1,k4_interva1(sK0,X1,X2)),X1) & m1_interva1(X2,sK0) & ~v1_xboole_0(X2)) & m1_interva1(X1,sK0) & ~v1_xboole_0(X1)) => (? [X2] : (~r1_interva1(sK0,k3_interva1(sK0,sK1,k4_interva1(sK0,sK1,X2)),sK1) & m1_interva1(X2,sK0) & ~v1_xboole_0(X2)) & m1_interva1(sK1,sK0) & ~v1_xboole_0(sK1)) [choice axiom]
108. ? [X2] : (~r1_interva1(sK0,k3_interva1(sK0,sK1,k4_interva1(sK0,sK1,X2)),sK1) & m1_interva1(X2,sK0) & ~v1_xboole_0(X2)) => (~r1_interva1(sK0,k3_interva1(sK0,sK1,k4_interva1(sK0,sK1,sK2)),sK1) & m1_interva1(sK2,sK0) & ~v1_xboole_0(sK2)) [choice axiom]
109. ((~r1_interva1(sK0,k3_interva1(sK0,sK1,k4_interva1(sK0,sK1,sK2)),sK1) & m1_interva1(sK2,sK0) & ~v1_xboole_0(sK2)) & m1_interva1(sK1,sK0) & ~v1_xboole_0(sK1)) & ~v1_xboole_0(sK0) [skolemisation 62,108,107,106]
126. ! [X0,X1,X2] : (((r1_interva1(X0,X1,X2) | X1 != X2) & (X1 = X2 | ~r1_interva1(X0,X1,X2))) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0)) [nnf transformation 105]
131. ~v1_xboole_0(sK0) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
132. ~v1_xboole_0(sK1) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
133. m1_interva1(sK1,sK0) [cnf transformation 109] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
134. ~v1_xboole_0(sK2) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
135. m1_interva1(sK2,sK0) [cnf transformation 109] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
136. ~r1_interva1(sK0,k3_interva1(sK0,sK1,k4_interva1(sK0,sK1,sK2)),sK1) [cnf transformation 109] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
144. ~m1_interva1(X1,X0) | v1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 68] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
146. k4_interva1(X0,X1,X2) = k2_setfam_1(X1,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 70] {a:0,w:20,wCS:8000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
147. k3_interva1(X0,X1,X2) = k3_setfam_1(X1,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 72] {a:0,w:20,wCS:8000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
158. m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~m1_interva1(X1,X0) [cnf transformation 80] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
160. k3_setfam_1(X1,k2_setfam_1(X1,X2)) = X1 | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(X1,X0) | v1_xboole_0(X1) [cnf transformation 84] {a:0,w:27,wCS:10800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
181. ~v1_xboole_0(k4_interva1(X0,X1,X2)) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 97] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
182. m1_interva1(k4_interva1(X0,X1,X2),X0) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 99] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
186. r1_interva1(X0,X1,X2) | X1 != X2 | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X1,X0) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 126] {a:0,w:19,wCS:7600,sine:1,thAx:0,allAx:1,thDist:-1}
189. r1_interva1(X0,X2,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | v1_xboole_0(X0) [equality resolution 186] {a:0,w:16,wCS:6400,sine:1,thAx:0,allAx:1,thDist:-1}
190. r1_interva1(X0,X2,X2) | ~m1_interva1(X2,X0) | v1_xboole_0(X2) | v1_xboole_0(X0) [duplicate literal removal 189] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
192. 1 <=> r1_interva1(sK0,k3_interva1(sK0,sK1,k4_interva1(sK0,sK1,sK2)),sK1) [avatar definition]
194. ~r1_interva1(sK0,k3_interva1(sK0,sK1,k4_interva1(sK0,sK1,sK2)),sK1) <- (~1) [avatar component clause 192] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
195. ~1 [avatar split clause 136,192]
197. 2 <=> m1_interva1(sK2,sK0) [avatar definition]
199. m1_interva1(sK2,sK0) <- (2) [avatar component clause 197] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
200. 2 [avatar split clause 135,197]
202. 3 <=> v1_xboole_0(sK2) [avatar definition]
204. ~v1_xboole_0(sK2) <- (~3) [avatar component clause 202] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
205. ~3 [avatar split clause 134,202]
207. 4 <=> m1_interva1(sK1,sK0) [avatar definition]
209. m1_interva1(sK1,sK0) <- (4) [avatar component clause 207] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
210. 4 [avatar split clause 133,207]
212. 5 <=> v1_xboole_0(sK1) [avatar definition]
214. ~v1_xboole_0(sK1) <- (~5) [avatar component clause 212] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
215. ~5 [avatar split clause 132,212]
217. 6 <=> v1_xboole_0(sK0) [avatar definition]
219. ~v1_xboole_0(sK0) <- (~6) [avatar component clause 217] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
220. ~6 [avatar split clause 131,217]
294. v1_interva1(sK2,sK0) | v1_xboole_0(sK2) | v1_xboole_0(sK0) <- (2) [resolution 144,199] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
295. v1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (4) [resolution 144,209] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
299. v1_interva1(sK2,sK0) | v1_xboole_0(sK0) <- (2, ~3) [subsumption resolution 294,204] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
300. v1_interva1(sK2,sK0) <- (2, ~3, ~6) [subsumption resolution 299,219] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
302. 12 <=> v1_interva1(sK2,sK0) [avatar definition]
304. v1_interva1(sK2,sK0) <- (12) [avatar component clause 302] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
305. 12 | ~2 | 3 | 6 [avatar split clause 300,217,202,197,302]
306. v1_interva1(sK1,sK0) | v1_xboole_0(sK0) <- (4, ~5) [subsumption resolution 295,214] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
307. v1_interva1(sK1,sK0) <- (4, ~5, ~6) [subsumption resolution 306,219] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
309. 13 <=> v1_interva1(sK1,sK0) [avatar definition]
311. v1_interva1(sK1,sK0) <- (13) [avatar component clause 309] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
312. 13 | ~4 | 5 | 6 [avatar split clause 307,217,212,207,309]
365. ~r1_interva1(sK0,k3_setfam_1(sK1,k4_interva1(sK0,sK1,sK2)),sK1) | ~m1_interva1(k4_interva1(sK0,sK1,sK2),sK0) | v1_xboole_0(k4_interva1(sK0,sK1,sK2)) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1) [superposition 194,147] {a:1,w:27,wCS:2700,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
371. ~r1_interva1(sK0,k3_setfam_1(sK1,k4_interva1(sK0,sK1,sK2)),sK1) | ~m1_interva1(k4_interva1(sK0,sK1,sK2),sK0) | v1_xboole_0(k4_interva1(sK0,sK1,sK2)) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1, 4) [subsumption resolution 365,209] {a:1,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
372. ~r1_interva1(sK0,k3_setfam_1(sK1,k4_interva1(sK0,sK1,sK2)),sK1) | ~m1_interva1(k4_interva1(sK0,sK1,sK2),sK0) | v1_xboole_0(k4_interva1(sK0,sK1,sK2)) | v1_xboole_0(sK0) <- (~1, 4, ~5) [subsumption resolution 371,214] {a:1,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
373. ~r1_interva1(sK0,k3_setfam_1(sK1,k4_interva1(sK0,sK1,sK2)),sK1) | ~m1_interva1(k4_interva1(sK0,sK1,sK2),sK0) | v1_xboole_0(k4_interva1(sK0,sK1,sK2)) <- (~1, 4, ~5, ~6) [subsumption resolution 372,219] {a:1,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
375. 15 <=> v1_xboole_0(k4_interva1(sK0,sK1,sK2)) [avatar definition]
377. v1_xboole_0(k4_interva1(sK0,sK1,sK2)) <- (15) [avatar component clause 375] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
379. 16 <=> m1_interva1(k4_interva1(sK0,sK1,sK2),sK0) [avatar definition]
381. ~m1_interva1(k4_interva1(sK0,sK1,sK2),sK0) <- (~16) [avatar component clause 379] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
383. 17 <=> r1_interva1(sK0,k3_setfam_1(sK1,k4_interva1(sK0,sK1,sK2)),sK1) [avatar definition]
385. ~r1_interva1(sK0,k3_setfam_1(sK1,k4_interva1(sK0,sK1,sK2)),sK1) <- (~17) [avatar component clause 383] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
386. 15 | ~16 | ~17 | 1 | ~4 | 5 | 6 [avatar split clause 373,217,212,207,192,383,379,375]
399. 20 <=> r1_interva1(sK0,k3_setfam_1(sK1,k2_setfam_1(sK1,sK2)),sK1) [avatar definition]
401. ~r1_interva1(sK0,k3_setfam_1(sK1,k2_setfam_1(sK1,sK2)),sK1) <- (~20) [avatar component clause 399] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
404. ~m1_interva1(sK2,sK0) | v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~16) [resolution 381,182] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
407. v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~16) [subsumption resolution 404,199] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
408. ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, ~16) [subsumption resolution 407,204] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
409. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, 4, ~16) [subsumption resolution 408,209] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
410. v1_xboole_0(sK0) <- (2, ~3, 4, ~5, ~16) [subsumption resolution 409,214] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
411. $false <- (2, ~3, 4, ~5, ~6, ~16) [subsumption resolution 410,219] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
412. ~2 | 3 | ~4 | 5 | 6 | 16 [avatar contradiction clause 411]
436. ~r1_interva1(sK0,k3_setfam_1(sK1,k2_setfam_1(sK1,sK2)),sK1) | ~m1_interva1(sK2,sK0) | v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~17) [superposition 385,146] {a:2,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
437. ~r1_interva1(sK0,k3_setfam_1(sK1,k2_setfam_1(sK1,sK2)),sK1) | v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~17) [subsumption resolution 436,199] {a:2,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
438. ~r1_interva1(sK0,k3_setfam_1(sK1,k2_setfam_1(sK1,sK2)),sK1) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, ~17) [subsumption resolution 437,204] {a:2,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
439. ~r1_interva1(sK0,k3_setfam_1(sK1,k2_setfam_1(sK1,sK2)),sK1) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, 4, ~17) [subsumption resolution 438,209] {a:2,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
440. ~r1_interva1(sK0,k3_setfam_1(sK1,k2_setfam_1(sK1,sK2)),sK1) | v1_xboole_0(sK0) <- (2, ~3, 4, ~5, ~17) [subsumption resolution 439,214] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
441. ~r1_interva1(sK0,k3_setfam_1(sK1,k2_setfam_1(sK1,sK2)),sK1) <- (2, ~3, 4, ~5, ~6, ~17) [subsumption resolution 440,219] {a:2,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
442. ~20 | ~2 | 3 | ~4 | 5 | 6 | 17 [avatar split clause 441,383,217,212,207,202,197,399]
461. ~r1_interva1(sK0,sK1,sK1) | ~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0) | v1_xboole_0(sK2) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | v1_xboole_0(sK1) <- (~20) [superposition 401,160] {a:3,w:24,wCS:2400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
462. ~r1_interva1(sK0,sK1,sK1) | ~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | v1_xboole_0(sK1) <- (~3, ~20) [subsumption resolution 461,204] {a:3,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
463. ~r1_interva1(sK0,sK1,sK1) | ~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) <- (~3, ~5, ~20) [subsumption resolution 462,214] {a:3,w:20,wCS:2000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
465. 23 <=> ! [X0] : (~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0)) [avatar definition]
466. ~m1_subset_1(sK2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK2,X0) <- (23) [avatar component clause 465] {a:3,w:16,wCS:1600,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
468. 24 <=> r1_interva1(sK0,sK1,sK1) [avatar definition]
470. ~r1_interva1(sK0,sK1,sK1) <- (~24) [avatar component clause 468] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
471. 23 | ~24 | 3 | 5 | 20 [avatar split clause 463,399,212,202,468,465]
479. ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~24) [resolution 470,190] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
481. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (4, ~24) [subsumption resolution 479,209] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
482. v1_xboole_0(sK0) <- (4, ~5, ~24) [subsumption resolution 481,214] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
483. $false <- (4, ~5, ~6, ~24) [subsumption resolution 482,219] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
484. ~4 | 5 | 6 | 24 [avatar contradiction clause 483]
493. ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~v1_interva1(sK1,X0) | ~v1_interva1(sK2,X0) | ~m1_interva1(sK2,X0) <- (23) [resolution 466,158] {a:4,w:14,wCS:1400,nSel:1,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
496. ~v1_interva1(sK2,X0) | ~v1_interva1(sK1,X0) | ~m1_interva1(sK2,X0) | ~m1_interva1(sK1,X0) <- (23) [resolution 493,158] {a:5,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
503. ~v1_interva1(sK1,sK0) | ~m1_interva1(sK2,sK0) | ~m1_interva1(sK1,sK0) <- (12, 23) [resolution 496,304] {a:6,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
505. ~m1_interva1(sK2,sK0) | ~m1_interva1(sK1,sK0) <- (12, 13, 23) [subsumption resolution 503,311] {a:6,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
506. ~m1_interva1(sK1,sK0) <- (2, 12, 13, 23) [subsumption resolution 505,199] {a:6,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
507. $false <- (2, 4, 12, 13, 23) [subsumption resolution 506,209] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
508. ~2 | ~4 | ~12 | ~13 | ~23 [avatar contradiction clause 507]
510. ~m1_interva1(sK2,sK0) | v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (15) [resolution 377,181] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
513. v1_xboole_0(sK2) | ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, 15) [subsumption resolution 510,199] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
514. ~m1_interva1(sK1,sK0) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, 15) [subsumption resolution 513,204] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
515. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2, ~3, 4, 15) [subsumption resolution 514,209] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
516. v1_xboole_0(sK0) <- (2, ~3, 4, ~5, 15) [subsumption resolution 515,214] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
517. $false <- (2, ~3, 4, ~5, ~6, 15) [subsumption resolution 516,219] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
518. ~2 | 3 | ~4 | 5 | 6 | ~15 [avatar contradiction clause 517]
519. $false [avatar sat refutation 195,200,205,210,215,220,305,312,386,412,442,471,484,508,518] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:303,thDist:-303}
% SZS output end Proof for interva1__t35_interva1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 48
% >>> Preprocessing
% Introduced skolems: 13
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 58
% Generated clauses: 290
% Active clauses: 94
% Passive clauses: 129
% Final active clauses: 86
% Final passive clauses: 29
% Main loop iterations started: 98
% 
% >>> Simplifying Inferences
% Duplicate literals: 41
% Fw subsumption resolutions: 68
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 12
% Equational tautologies: 2
% Forward subsumptions: 38
% 
% >>> Generating Inferences
% Binary resolution: 100
% Unit resulting resolution: 4
% Forward superposition: 8
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 12
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 32
% SAT solver unit clauses: 11
% 
% Memory used [KB]: 9722
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% splitting model update: 0.001 s ( own 0.0000 s ) 
% backward demodulation index maintenance: 0.001 s
% superposition: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

