i: [1,224,0,2,1,1,0,0] t136_xreal_1
a: [3,353,0,2,1,224]
s: 353
i: [1,223,0,2,1,1,0,0] t136_xreal_1
a: [3,358,0,2,1,223]
s: 358
i: [1,226,0,3,1,1,0,0] t136_xreal_1
a: [3,343,0,3,1,226]
s: 343
i: [1,225,0,3,1,1,0,0] t136_xreal_1
a: [3,348,0,3,1,225]
s: 348
i: [1,227,0,5,1,1,0,0] t136_xreal_1
a: [3,338,0,5,1,227]
s: 338
i: [1,228,0,2,1,0,0,255] fc1_numbers
a: [3,363,0,2,1,228]
s: 363
i: [1,229,0,2,1,0,0,2] fc1_xboole_0
a: [3,368,0,2,1,229]
s: 368
i: [1,236,0,2,1,0,0,3] fc6_ordinal1
a: [3,398,0,2,1,236]
s: 398
i: [1,235,0,2,1,0,0,3] fc6_ordinal1
a: [3,403,0,2,1,235]
s: 403
i: [1,302,0,2,1,0,0,2] rc2_xboole_0
a: [3,415,0,2,1,302]
s: 415
i: [1,303,0,2,1,0,0,4] rc4_ordinal1
a: [3,420,0,2,1,303]
s: 420
i: [1,304,0,2,1,0,0,4] rc1_ordinal1
a: [3,425,0,2,1,304]
s: 425
i: [1,305,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,430,0,2,1,305]
s: 430
i: [1,306,0,2,1,0,0,1] rc1_xreal_0
a: [3,435,0,2,1,306]
s: 435
i: [1,307,0,2,1,0,0,2] rc1_xxreal_0
a: [3,440,0,2,1,307]
s: 440
i: [1,308,0,2,1,0,0,2] rc1_xboole_0
a: [3,445,0,2,1,308]
s: 445
i: [1,312,0,2,1,0,0,255] rc2_ordinal1
a: [3,450,0,2,1,312]
s: 450
i: [1,311,0,2,1,0,0,255] rc2_ordinal1
a: [3,455,0,2,1,311]
s: 455
i: [1,310,0,2,1,0,0,255] rc2_ordinal1
a: [3,460,0,2,1,310]
s: 460
i: [1,309,0,2,1,0,0,255] rc2_ordinal1
a: [3,465,0,2,1,309]
s: 465
i: [1,314,0,2,1,0,0,4] rc5_ordinal1
a: [3,470,0,2,1,314]
s: 470
i: [1,313,0,2,1,0,0,4] rc5_ordinal1
a: [3,475,0,2,1,313]
s: 475
i: [1,316,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,480,0,2,1,316]
s: 480
i: [1,315,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,485,0,2,1,315]
s: 485
i: [1,320,0,2,1,0,0,2] rc2_xreal_0
a: [3,490,0,2,1,320]
s: 490
i: [1,319,0,2,1,0,0,2] rc2_xreal_0
a: [3,495,0,2,1,319]
s: 495
i: [1,318,0,2,1,0,0,2] rc2_xreal_0
a: [3,500,0,2,1,318]
s: 500
i: [1,317,0,2,1,0,0,2] rc2_xreal_0
a: [3,505,0,2,1,317]
s: 505
i: [1,324,0,2,1,0,0,2] rc3_xreal_0
a: [3,510,0,2,1,324]
s: 510
i: [1,323,0,2,1,0,0,2] rc3_xreal_0
a: [3,515,0,2,1,323]
s: 515
i: [1,322,0,2,1,0,0,2] rc3_xreal_0
a: [3,520,0,2,1,322]
s: 520
i: [1,321,0,2,1,0,0,2] rc3_xreal_0
a: [3,525,0,2,1,321]
s: 525
i: [1,326,0,2,1,0,0,2] rc2_xxreal_0
a: [3,530,0,2,1,326]
s: 530
i: [1,325,0,2,1,0,0,2] rc2_xxreal_0
a: [3,535,0,2,1,325]
s: 535
i: [1,328,0,2,1,0,0,2] rc3_xxreal_0
a: [3,540,0,2,1,328]
s: 540
i: [1,327,0,2,1,0,0,2] rc3_xxreal_0
a: [3,545,0,2,1,327]
s: 545
i: [1,332,0,2,1,0,0,2] rc4_xreal_0
a: [3,550,0,2,1,332]
s: 550
i: [1,331,0,2,1,0,0,2] rc4_xreal_0
a: [3,555,0,2,1,331]
s: 555
i: [1,330,0,2,1,0,0,2] rc4_xreal_0
a: [3,560,0,2,1,330]
s: 560
i: [1,329,0,2,1,0,0,2] rc4_xreal_0
a: [3,565,0,2,1,329]
s: 565
i: [1,334,0,2,1,0,0,2] rc4_xxreal_0
a: [3,570,0,2,1,334]
s: 570
i: [1,333,0,2,1,0,0,2] rc4_xxreal_0
a: [3,575,0,2,1,333]
s: 575
i: [1,287,0,3,1,0,0,255] reflexivity_r1_tarski
s: 287
i: [1,230,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,373,0,3,1,230]
s: 373
i: [1,231,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,378,0,3,1,231]
s: 378
r: 368
d: [2,577,0,2,1,2,23,368,378]
a: [3,581,0,2,1,577]
s: 581
i: [1,232,0,3,1,0,0,2] t1_numerals
a: [3,383,0,3,1,232]
s: 383
d: [2,583,0,3,1,2,22,383,378]
a: [3,587,0,3,1,583]
s: 587
i: [1,238,0,4,2,0,0,4] cc6_ordinal1
s: 238
i: [1,239,0,4,2,0,0,4] cc2_xcmplx_0
s: 239
i: [1,240,0,4,2,0,0,4] cc2_xreal_0
s: 240
d: [2,589,1,2,1,1,40,238,420]
a: [3,594,1,2,1,589]
s: 594
i: [1,241,0,4,2,0,0,4] cc2_xxreal_0
s: 241
d: [2,590,1,2,1,1,40,238,470]
a: [3,599,1,2,1,590]
s: 599
i: [1,242,0,4,2,0,0,255] cc1_ordinal1
s: 242
i: [1,243,0,4,2,0,0,255] cc1_ordinal1
s: 243
d: [2,601,1,2,1,1,40,239,420]
a: [3,606,1,2,1,601]
s: 606
i: [1,246,0,4,2,0,0,2] cc3_xreal_0
s: 246
d: [2,687,1,2,1,1,40,246,358]
a: [3,696,1,2,1,687]
s: 696
i: [1,247,0,4,2,0,0,2] cc4_xreal_0
s: 247
i: [1,256,0,4,2,0,0,255] cc9_ordinal1
s: 256
d: [2,688,1,2,1,1,40,246,353]
a: [3,701,1,2,1,688]
s: 701
i: [1,257,0,4,2,0,0,4] cc7_ordinal1
s: 257
d: [2,708,1,2,1,1,40,247,358]
a: [3,717,1,2,1,708]
s: 717
i: [1,258,0,4,2,0,0,4] cc3_ordinal1
s: 258
i: [1,286,0,4,1,0,0,3] existence_m1_subset_1
s: 286
d: [2,709,1,2,1,1,40,247,353]
a: [3,722,1,2,1,709]
s: 722
i: [1,233,0,4,1,0,0,255] dt_k5_numbers
a: [3,388,0,4,1,233]
s: 388
d: [2,602,1,2,1,1,40,239,470]
a: [3,611,1,2,1,602]
s: 611
i: [1,234,0,4,1,0,0,255] dt_k6_numbers
a: [3,393,0,4,1,234]
s: 393
d: [2,801,0,4,1,2,22,388,373]
a: [3,805,0,4,1,801]
s: 805
d: [2,613,1,2,1,1,40,240,420]
a: [3,618,1,2,1,613]
s: 618
d: [2,807,0,4,1,2,22,393,373]
a: [3,811,0,4,1,807]
s: 811
d: [2,614,1,2,1,1,40,240,470]
a: [3,623,1,2,1,614]
s: 623
i: [1,259,0,5,2,0,0,2] t6_boole
s: 259
i: [1,260,0,5,2,0,0,3] cc8_ordinal1
s: 260
d: [2,625,1,2,1,1,40,241,420]
a: [3,630,1,2,1,625]
s: 630
i: [1,261,0,5,2,0,0,255] cc1_xcmplx_0
s: 261
d: [2,626,1,2,1,1,40,241,470]
a: [3,635,1,2,1,626]
s: 635
i: [1,262,0,5,2,0,0,255] cc1_xreal_0
s: 262
i: [1,283,0,5,2,0,0,255] t7_boole
s: 283
d: [2,637,1,2,1,1,40,242,398]
a: [3,645,1,2,1,637]
s: 645
i: [1,300,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,410,0,5,2,300]
s: 410
d: [2,639,1,2,1,1,40,242,425]
a: [3,655,1,2,1,639]
s: 655
d: [2,819,0,5,2,1,22,259,378]
s: 819
i: [1,267,0,6,3,0,0,255] cc2_ordinal1
s: 267
d: [2,842,1,5,2,2,44,343,819]
s: 842
i: [1,270,0,6,3,0,0,2] cc5_xxreal_0
s: 270
d: [2,843,1,5,2,2,44,348,819]
s: 843
i: [1,272,0,6,3,0,0,2] cc5_xxreal_0
s: 272
i: [1,273,0,6,3,0,0,2] cc3_xxreal_0
s: 273
d: [2,841,1,7,2,2,44,338,819]
s: 841
i: [1,275,0,6,3,0,0,2] cc3_xxreal_0
s: 275
d: [2,890,2,8,2,3,40,841,410]
a: [3,894,2,4,1,890]
s: 894
i: [1,277,0,6,3,0,0,2] cc7_xxreal_0
s: 277
i: [1,278,0,6,3,0,0,2] cc7_xxreal_0
s: 278
d: [2,662,1,2,1,1,40,243,398]
a: [3,670,1,2,1,662]
s: 670
i: [1,279,0,6,2,0,0,255] t1_subset
s: 279
d: [2,664,1,2,1,1,40,243,425]
a: [3,680,1,2,1,664]
s: 680
i: [1,288,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 288
i: [1,244,0,7,3,0,0,4] cc5_ordinal1
s: 244
d: [2,689,1,2,1,1,40,246,435]
a: [3,706,1,2,1,689]
s: 706
i: [1,245,0,7,2,0,0,1] t5_arithm
s: 245
d: [2,710,1,2,1,1,40,247,435]
a: [3,727,1,2,1,710]
s: 727
i: [1,281,0,7,2,0,0,255] t3_subset
s: 281
i: [1,282,0,7,3,0,0,2] t8_boole
s: 282
d: [2,729,1,2,1,1,40,256,581]
a: [3,736,1,2,1,729]
s: 736
i: [1,237,0,8,3,0,0,255] cc10_ordinal1
s: 237
d: [2,730,1,2,1,1,40,256,445]
a: [3,741,1,2,1,730]
s: 741
i: [1,264,0,8,4,0,0,2] cc4_xxreal_0
s: 264
i: [1,266,0,8,4,0,0,2] cc6_xxreal_0
s: 266
d: [2,922,2,6,3,1,40,264,717]
a: [3,937,2,2,1,922]
s: 937
r: 937
i: [1,268,0,8,4,0,0,2] cc8_xxreal_0
s: 268
a: [3,936,2,2,1,922]
s: 936
i: [1,280,0,8,3,0,0,255] t2_subset
s: 280
i: [1,298,0,8,3,0,0,1] fc5_xcmplx_0
s: 298
a: [3,941,2,2,1,922]
s: 941
i: [1,299,0,8,3,0,0,1] fc8_xreal_0
s: 299
d: [2,923,2,6,3,1,40,264,722]
a: [3,950,2,2,1,923]
s: 950
r: 950
i: [1,285,0,9,3,0,0,255] t5_subset
s: 285
i: [1,284,0,10,3,0,0,255] t4_subset
s: 284
a: [3,949,2,2,1,923]
s: 949
i: [1,301,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 301
a: [3,954,2,2,1,923]
s: 954
i: [1,250,0,11,5,0,0,1] t2_real
s: 250
i: [1,251,0,11,5,0,0,1] t1_real
s: 251
a: [3,944,2,2,1,922]
s: 944
i: [1,252,0,11,5,0,0,1] t8_real
s: 252
a: [3,957,2,2,1,923]
s: 957
i: [1,253,0,11,5,0,0,1] t7_real
s: 253
i: [1,254,0,11,5,0,0,1] t4_real
s: 254
d: [2,1118,1,12,4,1,40,252,338]
a: [3,1127,1,4,1,1118]
s: 1127
d: [2,1141,2,0,0,3,52,358,353,1127,299]
e: 1141
d: [2,1142,2,4,2,1,40,1127,299]
d: [2,1144,2,2,1,2,21,1142,353]
d: [2,1145,2,0,0,3,21,1144,358]
e: 1145
r: 1127
i: [1,255,0,11,5,0,0,1] t3_real
s: 255
a: [3,1126,1,4,1,1118]
s: 1126
d: [2,1147,3,0,0,2,52,894,1126,247]
e: 1147
d: [2,1148,2,4,1,1,40,1126,247]
d: [2,1151,2,0,0,2,21,1148,894]
e: 1151
r: 894
i: [1,289,0,12,5,0,0,1] fc31_xreal_0
s: 289
i: [1,290,0,12,5,0,0,1] fc32_xreal_0
s: 290
d: [2,1081,1,8,4,1,40,250,343]
d: [2,1088,1,6,3,2,21,1081,353]
a: [3,1092,1,2,1,1088]
s: 1092
i: [1,291,0,12,5,0,0,1] fc34_xreal_0
s: 291
d: [2,1076,1,11,3,1,40,301,338]
a: [3,1161,1,2,1,1076]
s: 1161
i: [1,292,0,12,5,0,0,1] fc33_xreal_0
s: 292
i: [1,297,0,12,5,0,0,1] fc9_xcmplx_0
s: 297
a: [3,1135,1,4,1,1118]
s: 1135
d: [2,1185,3,0,0,5,52,944,358,957,353,1135,292]
e: 1185
d: [2,1186,2,8,4,1,40,1135,292]
d: [2,1191,2,6,3,2,21,1186,353]
d: [2,1192,2,4,2,3,21,1191,957]
d: [2,1193,2,2,1,4,21,1192,358]
d: [2,1194,2,0,0,5,21,1193,944]
e: 1194
r: 944
i: [1,248,0,13,6,0,0,1] t6_real
s: 248
a: [3,945,2,2,1,922]
s: 945
d: [2,1219,3,0,0,3,52,717,941,945,272]
e: 1219
d: [2,1220,3,4,2,1,40,945,272]
d: [2,1223,3,2,1,2,21,1220,941]
d: [2,1224,3,0,0,3,21,1223,717]
e: 1224
r: 941
i: [1,249,0,13,6,0,0,1] t5_real
s: 249
i: [1,296,0,14,4,0,0,255] existence_m2_subset_1
s: 296
a: [3,940,2,2,1,922]
s: 940
i: [1,293,0,15,5,0,0,255] dt_m2_subset_1
s: 293
a: [3,893,2,4,1,890]
s: 893
i: [1,294,0,15,5,0,0,255] redefinition_m2_subset_1
s: 294
i: [1,295,0,15,5,0,0,255] redefinition_m2_subset_1
s: 295
a: [3,898,2,4,1,890]
s: 898
d: [2,731,1,2,1,1,40,256,565]
a: [3,746,1,2,1,731]
s: 746
d: [2,1149,2,4,1,1,40,1126,246]
a: [3,1170,2,4,1,1149]
s: 1170
d: [2,732,1,2,1,1,40,256,575]
a: [3,751,1,2,1,732]
s: 751
d: [2,753,1,2,1,1,40,257,581]
a: [3,760,1,2,1,753]
s: 760
d: [2,1266,2,0,0,2,52,1161,760,241]
e: 1266
d: [2,1267,2,0,0,2,52,1092,760,240]
e: 1267
d: [2,1268,2,2,1,1,40,760,241]
d: [2,1275,2,0,0,2,21,1268,1161]
e: 1275
d: [2,1269,2,2,1,1,40,760,240]
d: [2,1277,2,0,0,2,21,1269,1092]
e: 1277
r: 1092
r: 1161
a: [3,1091,1,2,1,1088]
s: 1091
a: [3,1096,1,2,1,1088]
s: 1096
r: 940
r: 945
r: 1096
r: 1135
a: [3,1095,1,2,1,1088]
s: 1095
a: [3,1131,1,2,1,1118]
s: 1131
d: [2,820,1,2,1,1,40,260,587]
d: [2,1082,1,8,4,1,40,250,348]
d: [2,1098,1,6,3,2,21,1082,358]
d: [2,1302,2,4,2,1,40,1096,270]
d: [2,1305,2,2,1,2,21,1302,581]
d: [2,1316,2,4,2,1,40,1131,273]
d: [2,1318,2,2,1,2,21,1316,581]
f: 1318,1305,1268,1269,1193,1145,1118,1098,1088,820,583,577,232,231,229,223,224,225,226,227
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t136_xreal_1
% SZS output start Proof for xreal_1__t136_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0)) => r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1))))) [input t136_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0)) => r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1))))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
12. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
15. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
21. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc33_xreal_0]
32. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
40. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
51. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
60. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
73. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
82. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
87. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
100. ? [X0] : (? [X1] : ((~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1)) & (r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
101. ? [X0] : (? [X1] : (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1)) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 100]
105. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
106. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
116. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 32]
117. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 116]
120. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 40]
121. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 120]
132. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 51]
143. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 15]
144. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 143]
145. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 12]
146. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 145]
165. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 21]
166. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 165]
177. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 87]
178. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 177]
183. ? [X0] : (? [X1] : (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(X0,X1)) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,X1)) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
184. ? [X1] : (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,X1)) & r1_xxreal_0(k6_numbers,X1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) => (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) & r1_xxreal_0(k6_numbers,sK1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) [choice axiom]
185. (~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) & r1_xxreal_0(k6_numbers,sK1) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 101,184,183]
223. v1_xreal_0(sK0) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
224. v1_xreal_0(sK1) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
225. r1_xxreal_0(k6_numbers,sK0) [cnf transformation 185] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
226. r1_xxreal_0(k6_numbers,sK1) [cnf transformation 185] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
227. ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) [cnf transformation 185] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
229. v1_xboole_0(k1_xboole_0) [cnf transformation 73] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
231. k6_numbers = k1_xboole_0 [cnf transformation 82] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
232. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 60] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
240. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 105] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
241. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 106] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
250. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 117] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
252. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 121] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
260. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 132] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
270. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 144] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
273. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 146] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
292. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 166] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
299. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 178] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
336. 1 <=> r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) [avatar definition]
338. ~r1_xxreal_0(k6_numbers,k7_xcmplx_0(sK0,sK1)) <- (~1) [avatar component clause 336] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
339. ~1 [avatar split clause 227,336]
341. 2 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
343. r1_xxreal_0(k6_numbers,sK1) <- (2) [avatar component clause 341] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
344. 2 [avatar split clause 226,341]
346. 3 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
348. r1_xxreal_0(k6_numbers,sK0) <- (3) [avatar component clause 346] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
349. 3 [avatar split clause 225,346]
351. 4 <=> v1_xreal_0(sK1) [avatar definition]
353. v1_xreal_0(sK1) <- (4) [avatar component clause 351] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
354. 4 [avatar split clause 224,351]
356. 5 <=> v1_xreal_0(sK0) [avatar definition]
358. v1_xreal_0(sK0) <- (5) [avatar component clause 356] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
359. 5 [avatar split clause 223,356]
366. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
368. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 366] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
369. 7 [avatar split clause 229,366]
376. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
378. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 376] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
379. 9 [avatar split clause 231,376]
381. 10 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
383. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (10) [avatar component clause 381] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
384. 10 [avatar split clause 232,381]
577. v1_xboole_0(k6_numbers) <- (7, 9) [backward demodulation 368,378] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
579. 50 <=> v1_xboole_0(k6_numbers) [avatar definition]
581. v1_xboole_0(k6_numbers) <- (50) [avatar component clause 579] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
582. 50 | ~7 | ~9 [avatar split clause 577,376,366,579]
583. m1_subset_1(k6_numbers,k4_ordinal1) <- (9, 10) [forward demodulation 383,378] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
585. 51 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
587. m1_subset_1(k6_numbers,k4_ordinal1) <- (51) [avatar component clause 585] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
588. 51 | ~9 | ~10 [avatar split clause 583,381,376,585]
758. 78 <=> v7_ordinal1(k6_numbers) [avatar definition]
760. v7_ordinal1(k6_numbers) <- (78) [avatar component clause 758] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
820. v7_ordinal1(k6_numbers) <- (51) [resolution 260,587] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
822. 78 | ~51 [avatar split clause 820,585,758]
943. 103 <=> v3_xxreal_0(sK0) [avatar definition]
956. 106 <=> v3_xxreal_0(sK1) [avatar definition]
957. ~v3_xxreal_0(sK1) <- (~106) [avatar component clause 956] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1081. ~v3_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (2) [resolution 250,343] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1082. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (3) [resolution 250,348] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1088. ~v3_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 4) [subsumption resolution 1081,353] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1090. 127 <=> v1_xreal_0(k6_numbers) [avatar definition]
1094. 128 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1096. v3_xxreal_0(k6_numbers) <- (128) [avatar component clause 1094] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1097. ~127 | 128 | ~106 | ~2 | ~4 [avatar split clause 1088,351,341,956,1094,1090]
1098. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 5) [subsumption resolution 1082,358] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1099. ~127 | 128 | ~103 | ~3 | ~5 [avatar split clause 1098,356,346,943,1094,1090]
1118. v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (~1) [resolution 252,338] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1125. 130 <=> v1_xreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1127. ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~130) [avatar component clause 1125] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1129. 131 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1131. v2_xxreal_0(k6_numbers) <- (131) [avatar component clause 1129] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1133. 132 <=> v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1135. v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) <- (132) [avatar component clause 1133] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1136. ~127 | ~130 | 131 | 132 | 1 [avatar split clause 1118,336,1133,1129,1125,1090]
1142. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~130) [resolution 1127,299] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1144. ~v1_xreal_0(sK0) <- (4, ~130) [subsumption resolution 1142,353] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1145. $false <- (4, 5, ~130) [subsumption resolution 1144,358] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1146. ~4 | ~5 | 130 [avatar contradiction clause 1145]
1159. 135 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1186. ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (132) [resolution 1135,292] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1191. v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (4, 132) [subsumption resolution 1186,353] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1192. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (4, ~106, 132) [subsumption resolution 1191,957] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1193. v3_xxreal_0(sK0) <- (4, 5, ~106, 132) [subsumption resolution 1192,358] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1196. 103 | ~4 | ~5 | 106 | ~132 [avatar split clause 1193,1133,956,356,351,943]
1268. v1_xxreal_0(k6_numbers) <- (78) [resolution 760,241] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1269. v1_xreal_0(k6_numbers) <- (78) [resolution 760,240] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1279. 127 | ~78 [avatar split clause 1269,758,1090]
1289. 135 | ~78 [avatar split clause 1268,758,1159]
1302. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (128) [resolution 1096,270] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1305. ~v1_xxreal_0(k6_numbers) <- (50, 128) [subsumption resolution 1302,581] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1306. ~135 | ~50 | ~128 [avatar split clause 1305,1094,579,1159]
1316. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (131) [resolution 1131,273] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1318. ~v1_xxreal_0(k6_numbers) <- (50, 131) [subsumption resolution 1316,581] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1319. ~135 | ~50 | ~131 [avatar split clause 1318,1129,579,1159]
1320. $false [avatar sat refutation 339,344,349,354,359,369,379,384,582,588,822,1097,1099,1136,1146,1196,1279,1289,1306,1319] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:895,thDist:-895}
% SZS output end Proof for xreal_1__t136_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 87
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 112
% Generated clauses: 530
% Active clauses: 148
% Passive clauses: 233
% Final active clauses: 135
% Final passive clauses: 68
% Main loop iterations started: 160
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 51
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 18
% Forward subsumptions: 56
% 
% >>> Generating Inferences
% Binary resolution: 157
% Unit resulting resolution: 6
% Factoring: 1
% Forward superposition: 16
% Backward superposition: 6
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 32
% Split components: 79
% Sat splitting refutations: 12
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 164
% SAT solver unit clauses: 50
% SAT solver binary clauses: 60
% 
% Memory used [KB]: 10106
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% splitting model update: 0.001 s
% passive container maintenance: 0.002 s
% other: 0.004 s ( own 0.0000 s ) 

