i: [1,93,0,2,1,1,0,0] t34_yellow_5
a: [3,160,0,2,1,93]
s: 160
i: [1,92,0,2,1,1,0,0] t34_yellow_5
a: [3,165,0,2,1,92]
s: 165
i: [1,91,0,2,1,1,0,0] t34_yellow_5
a: [3,170,0,2,1,91]
s: 170
i: [1,94,0,4,1,1,0,0] t34_yellow_5
a: [3,155,0,4,1,94]
s: 155
i: [1,141,0,2,1,0,0,1] existence_l1_orders_2
a: [3,175,0,2,1,141]
s: 175
i: [1,95,0,18,2,1,0,0] t34_yellow_5
a: [3,146,0,9,1,95]
s: 146
i: [1,131,0,4,1,0,0,1] existence_m1_subset_1
s: 131
i: [1,98,0,6,3,0,0,255] cc4_yellow_0
s: 98
i: [1,99,0,6,3,0,0,255] cc4_yellow_0
s: 99
i: [1,100,0,6,3,0,0,2] cc1_lattice3
s: 100
i: [1,101,0,6,3,0,0,2] cc2_lattice3
s: 101
i: [1,96,0,7,2,0,0,1] dt_k3_yellow_0
s: 96
i: [1,97,0,7,2,0,0,1] dt_k4_yellow_0
s: 97
i: [1,103,0,8,4,0,0,255] cc6_waybel_1
s: 103
i: [1,105,0,8,4,0,0,255] cc7_waybel_1
s: 105
i: [1,107,0,8,4,0,0,255] cc5_waybel_1
s: 107
i: [1,108,0,8,4,0,0,255] cc5_waybel_1
s: 108
i: [1,109,0,8,4,0,0,255] cc5_waybel_1
s: 109
i: [1,110,0,8,4,0,0,255] cc5_waybel_1
s: 110
i: [1,111,0,8,4,0,0,255] cc5_waybel_1
s: 111
i: [1,115,0,8,4,0,0,255] cc10_waybel_1
s: 115
d: [2,177,1,6,3,1,40,115,160]
d: [2,179,1,4,2,2,21,177,165]
d: [2,180,1,2,1,3,21,179,170]
a: [3,184,1,2,1,180]
s: 184
i: [1,117,0,8,4,0,0,255] cc8_waybel_1
s: 117
d: [2,199,2,6,3,1,40,184,103]
d: [2,206,2,4,2,2,21,199,170]
d: [2,207,2,2,1,3,21,206,160]
a: [3,211,2,2,1,207]
s: 211
i: [1,118,0,8,4,0,0,255] cc8_waybel_1
s: 118
i: [1,119,0,8,4,0,0,255] cc8_waybel_1
s: 119
d: [2,200,2,6,3,1,40,184,105]
d: [2,213,2,4,2,2,21,200,170]
d: [2,214,2,2,1,3,21,213,160]
a: [3,218,2,2,1,214]
s: 218
i: [1,120,0,8,4,0,0,255] cc8_waybel_1
s: 120
d: [2,201,2,6,3,1,40,184,107]
d: [2,220,2,4,2,2,21,201,170]
d: [2,221,2,2,1,3,21,220,160]
a: [3,225,2,2,1,221]
s: 225
i: [1,121,0,8,4,0,0,255] cc8_waybel_1
s: 121
i: [1,122,0,8,4,0,0,255] cc8_waybel_1
s: 122
d: [2,305,1,6,3,1,40,122,160]
d: [2,307,1,4,2,2,21,305,165]
d: [2,308,1,2,1,3,21,307,170]
a: [3,312,1,2,1,308]
s: 312
i: [1,123,0,8,4,0,0,255] cc8_waybel_1
s: 123
d: [2,202,2,6,3,1,40,184,108]
d: [2,227,2,4,2,2,21,202,170]
d: [2,228,2,2,1,3,21,227,160]
a: [3,232,2,2,1,228]
s: 232
i: [1,124,0,8,4,0,0,255] cc8_waybel_1
s: 124
i: [1,125,0,8,4,0,0,255] cc5_yellow_0
s: 125
d: [2,334,1,6,3,1,40,124,160]
d: [2,336,1,4,2,2,21,334,165]
d: [2,337,1,2,1,3,21,336,170]
a: [3,341,1,2,1,337]
s: 341
i: [1,132,0,14,4,0,0,1] dt_k7_waybel_1
s: 132
d: [2,203,2,6,3,1,40,184,109]
d: [2,234,2,4,2,2,21,203,170]
d: [2,235,2,2,1,3,21,234,160]
a: [3,239,2,2,1,235]
s: 239
i: [1,139,0,17,4,0,0,255] dt_k10_lattice3
s: 139
i: [1,140,0,17,4,0,0,255] dt_k11_lattice3
s: 140
d: [2,204,2,6,3,1,40,184,110]
d: [2,241,2,4,2,2,21,204,170]
d: [2,242,2,2,1,3,21,241,160]
a: [3,246,2,2,1,242]
s: 246
i: [1,133,0,21,6,0,0,1] dt_k12_lattice3
s: 133
d: [2,205,2,6,3,1,40,184,111]
d: [2,248,2,4,2,2,21,205,170]
d: [2,249,2,2,1,3,21,248,160]
a: [3,253,2,2,1,249]
s: 253
i: [1,134,0,21,6,0,0,1] dt_k13_lattice3
s: 134
i: [1,113,0,22,11,0,0,255] cc9_waybel_1
s: 113
d: [2,320,2,4,2,1,40,312,98]
d: [2,321,2,2,1,2,21,320,160]
a: [3,325,2,2,1,321]
s: 325
d: [2,351,0,20,10,0,21,113,100]
s: 351
d: [2,178,1,6,3,1,40,115,175]
a: [3,197,1,2,1,178]
s: 197
i: [1,126,0,23,6,0,0,255] d23_waybel_1
s: 126
i: [1,127,0,23,6,0,0,255] d23_waybel_1
s: 127
i: [1,135,0,23,6,0,0,255] redefinition_k12_lattice3
s: 135
i: [1,136,0,23,6,0,0,255] redefinition_k13_lattice3
s: 136
i: [1,137,0,23,6,0,0,1] commutativity_k12_lattice3
s: 137
i: [1,138,0,23,6,0,0,1] commutativity_k13_lattice3
s: 138
i: [1,128,0,30,7,0,0,255] d23_waybel_1
s: 128
d: [2,357,1,37,11,0,44,135,127]
d: [2,362,1,27,8,0,18,357]
d: [2,363,1,25,7,0,21,362,101]
s: 363
i: [1,142,0,30,10,0,0,255] t11_yellow_5
s: 142
d: [2,388,2,27,7,1,44,146,363]
d: [2,395,2,22,6,1,19,388]
d: [2,398,2,18,5,2,21,395,155]
d: [2,399,2,16,4,3,21,398,160]
d: [2,400,2,14,3,4,21,399,253]
d: [2,401,2,12,2,5,21,400,239]
a: [3,405,2,6,1,401]
s: 405
i: [1,129,0,31,11,0,0,255] t11_yellow_5
s: 129
d: [2,416,1,33,11,1,44,146,129]
d: [2,431,1,29,10,2,21,416,155]
d: [2,432,1,27,9,3,21,431,160]
d: [2,433,1,25,8,4,21,432,165]
d: [2,434,1,23,7,5,21,433,253]
d: [2,435,1,21,6,6,21,434,246]
d: [2,436,1,19,5,7,21,435,239]
d: [2,437,1,17,4,8,21,436,232]
d: [2,438,1,15,3,9,21,437,225]
s: 438
d: [2,439,2,31,9,9,44,438,363]
d: [2,444,2,23,7,9,18,439]
d: [2,445,2,18,6,9,19,444]
d: [2,446,2,14,5,9,21,445,155]
d: [2,447,2,12,4,9,21,446,160]
d: [2,448,2,10,3,9,21,447,253]
d: [2,449,2,8,2,9,21,448,239]
s: 449
r: 438
d: [2,411,1,32,12,0,40,142,132]
d: [2,412,1,26,10,0,18,411]
d: [2,413,1,24,9,0,21,412,100]
s: 413
r: 142
d: [2,458,3,0,0,9,52,225,232,239,246,253,165,160,155,405,413]
e: 458
d: [2,459,3,18,8,1,40,413,405]
d: [2,464,3,14,7,2,21,459,155]
d: [2,465,3,12,6,3,21,464,160]
d: [2,466,3,10,5,4,21,465,165]
d: [2,467,3,8,4,5,21,466,253]
d: [2,468,3,6,3,6,21,467,246]
d: [2,469,3,4,2,7,21,468,239]
d: [2,470,3,2,1,8,21,469,232]
d: [2,471,3,0,0,9,21,470,225]
e: 471
r: 405
a: [3,404,2,6,1,401]
s: 404
d: [2,358,1,37,11,0,44,127,135]
d: [2,361,1,27,8,0,18,358]
d: [2,364,1,25,7,0,21,361,101]
s: 364
d: [2,365,1,37,11,0,44,136,126]
d: [2,370,1,27,8,0,18,365]
d: [2,371,1,25,7,0,21,370,100]
s: 371
a: [3,409,2,6,1,401]
s: 409
d: [2,495,3,0,0,4,52,170,160,155,409,132]
e: 495
d: [2,496,3,8,3,1,40,409,132]
d: [2,500,3,4,2,2,21,496,155]
d: [2,501,3,2,1,3,21,500,160]
d: [2,502,3,0,0,4,21,501,170]
e: 502
r: 449
r: 146
r: 409
a: [3,145,0,9,1,95]
s: 145
a: [3,408,2,6,1,401]
s: 408
a: [3,150,0,9,1,95]
s: 150
d: [2,507,3,0,0,7,52,160,239,246,155,404,408,150,371]
e: 507
d: [2,509,2,27,7,1,44,150,371]
d: [2,510,2,22,6,1,19,509]
d: [2,520,2,16,5,2,21,510,408]
d: [2,521,2,12,4,3,21,520,155]
d: [2,522,2,10,3,4,21,521,160]
d: [2,523,2,8,2,5,21,522,246]
d: [2,524,2,6,1,6,21,523,239]
d: [2,525,2,0,0,7,21,524,404]
e: 525
d: [2,255,1,6,3,1,40,117,160]
d: [2,257,1,4,2,2,21,255,165]
d: [2,258,1,2,1,3,21,257,170]
d: [2,265,1,6,3,1,40,118,160]
d: [2,267,1,4,2,2,21,265,165]
d: [2,268,1,2,1,3,21,267,170]
d: [2,275,1,6,3,1,40,119,160]
d: [2,277,1,4,2,2,21,275,165]
d: [2,278,1,2,1,3,21,277,170]
d: [2,285,1,6,3,1,40,120,160]
d: [2,287,1,4,2,2,21,285,165]
d: [2,288,1,2,1,3,21,287,170]
d: [2,295,1,6,3,1,40,121,160]
d: [2,297,1,4,2,2,21,295,165]
d: [2,298,1,2,1,3,21,297,170]
d: [2,482,3,6,1,10,40,404,449]
f: 525,502,482,471,298,288,278,268,258,91,92,93,94,95
% Refutation found. Thanks to Tanya!
% SZS status Theorem for yellow_5__t34_yellow_5
% SZS output start Proof for yellow_5__t34_yellow_5
1. ! [X0] : ((l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => (k13_lattice3(X0,X1,k7_waybel_1(X0,X1)) = k4_yellow_0(X0) & k12_lattice3(X0,X1,k7_waybel_1(X0,X1)) = k3_yellow_0(X0)))) [input t34_yellow_5]
2. ~! [X0] : ((l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => (k13_lattice3(X0,X1,k7_waybel_1(X0,X1)) = k4_yellow_0(X0) & k12_lattice3(X0,X1,k7_waybel_1(X0,X1)) = k3_yellow_0(X0)))) [negated conjecture 1]
4. ! [X0] : (l1_orders_2(X0) => (v1_lattice3(X0) => ~v2_struct_0(X0))) [input cc1_lattice3]
5. ! [X0] : (l1_orders_2(X0) => (v2_lattice3(X0) => ~v2_struct_0(X0))) [input cc2_lattice3]
11. ! [X0] : (l1_orders_2(X0) => ((v11_waybel_1(X0) & ~v2_struct_0(X0)) => (v10_waybel_1(X0) & v2_waybel_1(X0) & v3_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & ~v2_struct_0(X0)))) [input cc8_waybel_1]
15. ! [X0] : ((l1_orders_2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r6_waybel_1(X0,X1,X2) <=> (k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2)))))) [input d23_waybel_1]
22. ! [X0,X1] : ((m1_subset_1(X1,u1_struct_0(X0)) & l1_orders_2(X0) & ~v2_struct_0(X0)) => m1_subset_1(k7_waybel_1(X0,X1),u1_struct_0(X0))) [input dt_k7_waybel_1]
30. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_orders_2(X0) & v2_lattice3(X0) & v5_orders_2(X0)) => k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2)) [input redefinition_k12_lattice3]
31. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_orders_2(X0) & v1_lattice3(X0) & v5_orders_2(X0)) => k13_lattice3(X0,X1,X2) = k10_lattice3(X0,X1,X2)) [input redefinition_k13_lattice3]
32. ! [X0] : ((l1_orders_2(X0) & v11_waybel_1(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r6_waybel_1(X0,X1,X2) <=> k7_waybel_1(X0,X1) = X2)))) [input t11_yellow_5]
35. ? [X0] : (? [X1] : ((k13_lattice3(X0,X1,k7_waybel_1(X0,X1)) != k4_yellow_0(X0) | k12_lattice3(X0,X1,k7_waybel_1(X0,X1)) != k3_yellow_0(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
36. ? [X0] : (? [X1] : ((k13_lattice3(X0,X1,k7_waybel_1(X0,X1)) != k4_yellow_0(X0) | k12_lattice3(X0,X1,k7_waybel_1(X0,X1)) != k3_yellow_0(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0)) [flattening 35]
41. ! [X0] : ((~v2_struct_0(X0) | ~v1_lattice3(X0)) | ~l1_orders_2(X0)) [ennf transformation 4]
42. ! [X0] : (~v2_struct_0(X0) | ~v1_lattice3(X0) | ~l1_orders_2(X0)) [flattening 41]
43. ! [X0] : ((~v2_struct_0(X0) | ~v2_lattice3(X0)) | ~l1_orders_2(X0)) [ennf transformation 5]
44. ! [X0] : (~v2_struct_0(X0) | ~v2_lattice3(X0) | ~l1_orders_2(X0)) [flattening 43]
55. ! [X0] : (((v10_waybel_1(X0) & v2_waybel_1(X0) & v3_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & ~v2_struct_0(X0)) | (~v11_waybel_1(X0) | v2_struct_0(X0))) | ~l1_orders_2(X0)) [ennf transformation 11]
56. ! [X0] : ((v10_waybel_1(X0) & v2_waybel_1(X0) & v3_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v5_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & ~v2_struct_0(X0)) | ~v11_waybel_1(X0) | v2_struct_0(X0) | ~l1_orders_2(X0)) [flattening 55]
59. ! [X0] : (! [X1] : (! [X2] : ((r6_waybel_1(X0,X1,X2) <=> (k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 15]
60. ! [X0] : (! [X1] : (! [X2] : ((r6_waybel_1(X0,X1,X2) <=> (k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [flattening 59]
61. ! [X0] : (! [X1] : (! [X2] : ((r6_waybel_1(X0,X1,X2) <=> k7_waybel_1(X0,X1) = X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0))) [ennf transformation 32]
62. ! [X0] : (! [X1] : (! [X2] : ((r6_waybel_1(X0,X1,X2) <=> k7_waybel_1(X0,X1) = X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0)) [flattening 61]
63. ! [X0,X1] : (m1_subset_1(k7_waybel_1(X0,X1),u1_struct_0(X0)) | (~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 22]
64. ! [X0,X1] : (m1_subset_1(k7_waybel_1(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0)) [flattening 63]
69. ! [X0,X1,X2] : (k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0))) [ennf transformation 30]
70. ! [X0,X1,X2] : (k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0)) [flattening 69]
71. ! [X0,X1,X2] : (k13_lattice3(X0,X1,X2) = k10_lattice3(X0,X1,X2) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0))) [ennf transformation 31]
72. ! [X0,X1,X2] : (k13_lattice3(X0,X1,X2) = k10_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0)) [flattening 71]
81. ? [X0] : (? [X1] : ((k13_lattice3(X0,X1,k7_waybel_1(X0,X1)) != k4_yellow_0(X0) | k12_lattice3(X0,X1,k7_waybel_1(X0,X1)) != k3_yellow_0(X0)) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v11_waybel_1(X0) & ~v2_struct_0(X0)) => (? [X1] : ((k13_lattice3(sK0,X1,k7_waybel_1(sK0,X1)) != k4_yellow_0(sK0) | k12_lattice3(sK0,X1,k7_waybel_1(sK0,X1)) != k3_yellow_0(sK0)) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v11_waybel_1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
82. ? [X1] : ((k13_lattice3(sK0,X1,k7_waybel_1(sK0,X1)) != k4_yellow_0(sK0) | k12_lattice3(sK0,X1,k7_waybel_1(sK0,X1)) != k3_yellow_0(sK0)) & m1_subset_1(X1,u1_struct_0(sK0))) => ((k4_yellow_0(sK0) != k13_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1)) | k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
83. ((k4_yellow_0(sK0) != k13_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1)) | k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v11_waybel_1(sK0) & ~v2_struct_0(sK0) [skolemisation 36,82,81]
84. ! [X0] : (! [X1] : (! [X2] : (((r6_waybel_1(X0,X1,X2) | (k3_yellow_0(X0) != k11_lattice3(X0,X1,X2) | k4_yellow_0(X0) != k10_lattice3(X0,X1,X2))) & ((k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2)) | ~r6_waybel_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [nnf transformation 60]
85. ! [X0] : (! [X1] : (! [X2] : (((r6_waybel_1(X0,X1,X2) | k3_yellow_0(X0) != k11_lattice3(X0,X1,X2) | k4_yellow_0(X0) != k10_lattice3(X0,X1,X2)) & ((k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) & k4_yellow_0(X0) = k10_lattice3(X0,X1,X2)) | ~r6_waybel_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | v2_struct_0(X0)) [flattening 84]
86. ! [X0] : (! [X1] : (! [X2] : (((r6_waybel_1(X0,X1,X2) | k7_waybel_1(X0,X1) != X2) & (k7_waybel_1(X0,X1) = X2 | ~r6_waybel_1(X0,X1,X2))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0)) [nnf transformation 62]
91. ~v2_struct_0(sK0) [cnf transformation 83] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
92. v11_waybel_1(sK0) [cnf transformation 83] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
93. l1_orders_2(sK0) [cnf transformation 83] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
94. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 83] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
95. k4_yellow_0(sK0) != k13_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1)) | k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1)) [cnf transformation 83] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
100. ~v1_lattice3(X0) | ~v2_struct_0(X0) | ~l1_orders_2(X0) [cnf transformation 42] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
101. ~v2_lattice3(X0) | ~v2_struct_0(X0) | ~l1_orders_2(X0) [cnf transformation 44] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
117. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v3_orders_2(X0) [cnf transformation 56] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
118. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v4_orders_2(X0) [cnf transformation 56] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
119. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v5_orders_2(X0) [cnf transformation 56] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
120. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v1_lattice3(X0) [cnf transformation 56] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
121. ~l1_orders_2(X0) | ~v11_waybel_1(X0) | v2_struct_0(X0) | v2_lattice3(X0) [cnf transformation 56] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
126. k4_yellow_0(X0) = k10_lattice3(X0,X1,X2) | ~r6_waybel_1(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0) [cnf transformation 85] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
127. k3_yellow_0(X0) = k11_lattice3(X0,X1,X2) | ~r6_waybel_1(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0) [cnf transformation 85] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
129. k7_waybel_1(X0,X1) = X2 | ~r6_waybel_1(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) [cnf transformation 86] {a:0,w:31,wCS:12400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
130. r6_waybel_1(X0,X1,X2) | k7_waybel_1(X0,X1) != X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) [cnf transformation 86] {a:0,w:31,wCS:12400,sine:255,thAx:0,allAx:1,thDist:-1}
132. m1_subset_1(k7_waybel_1(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0) [cnf transformation 64] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
135. k12_lattice3(X0,X1,X2) = k11_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v2_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 70] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
136. k13_lattice3(X0,X1,X2) = k10_lattice3(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 72] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
142. ~m1_subset_1(k7_waybel_1(X0,X1),u1_struct_0(X0)) | r6_waybel_1(X0,X1,k7_waybel_1(X0,X1)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) [equality resolution 130] {a:0,w:30,wCS:12000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
144. 1 <=> k3_yellow_0(sK0) = k12_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1)) [avatar definition]
146. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~1) [avatar component clause 144] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
148. 2 <=> k4_yellow_0(sK0) = k13_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1)) [avatar definition]
150. k4_yellow_0(sK0) != k13_lattice3(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~2) [avatar component clause 148] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
151. ~1 | ~2 [avatar split clause 95,148,144]
153. 3 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
155. m1_subset_1(sK1,u1_struct_0(sK0)) <- (3) [avatar component clause 153] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
156. 3 [avatar split clause 94,153]
158. 4 <=> l1_orders_2(sK0) [avatar definition]
160. l1_orders_2(sK0) <- (4) [avatar component clause 158] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
161. 4 [avatar split clause 93,158]
163. 5 <=> v11_waybel_1(sK0) [avatar definition]
165. v11_waybel_1(sK0) <- (5) [avatar component clause 163] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
166. 5 [avatar split clause 92,163]
168. 6 <=> v2_struct_0(sK0) [avatar definition]
170. ~v2_struct_0(sK0) <- (~6) [avatar component clause 168] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
171. ~6 [avatar split clause 91,168]
223. 14 <=> v3_orders_2(sK0) [avatar definition]
225. v3_orders_2(sK0) <- (14) [avatar component clause 223] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
230. 15 <=> v4_orders_2(sK0) [avatar definition]
232. v4_orders_2(sK0) <- (15) [avatar component clause 230] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
237. 16 <=> v5_orders_2(sK0) [avatar definition]
239. v5_orders_2(sK0) <- (16) [avatar component clause 237] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
244. 17 <=> v1_lattice3(sK0) [avatar definition]
246. v1_lattice3(sK0) <- (17) [avatar component clause 244] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
251. 18 <=> v2_lattice3(sK0) [avatar definition]
253. v2_lattice3(sK0) <- (18) [avatar component clause 251] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
255. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v3_orders_2(sK0) <- (4) [resolution 117,160] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
257. v2_struct_0(sK0) | v3_orders_2(sK0) <- (4, 5) [subsumption resolution 255,165] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
258. v3_orders_2(sK0) <- (4, 5, ~6) [subsumption resolution 257,170] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
259. 14 | ~4 | ~5 | 6 [avatar split clause 258,168,163,158,223]
265. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v4_orders_2(sK0) <- (4) [resolution 118,160] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
267. v2_struct_0(sK0) | v4_orders_2(sK0) <- (4, 5) [subsumption resolution 265,165] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
268. v4_orders_2(sK0) <- (4, 5, ~6) [subsumption resolution 267,170] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
269. 15 | ~4 | ~5 | 6 [avatar split clause 268,168,163,158,230]
275. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v5_orders_2(sK0) <- (4) [resolution 119,160] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
277. v2_struct_0(sK0) | v5_orders_2(sK0) <- (4, 5) [subsumption resolution 275,165] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
278. v5_orders_2(sK0) <- (4, 5, ~6) [subsumption resolution 277,170] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
279. 16 | ~4 | ~5 | 6 [avatar split clause 278,168,163,158,237]
285. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v1_lattice3(sK0) <- (4) [resolution 120,160] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
287. v2_struct_0(sK0) | v1_lattice3(sK0) <- (4, 5) [subsumption resolution 285,165] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
288. v1_lattice3(sK0) <- (4, 5, ~6) [subsumption resolution 287,170] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
289. 17 | ~4 | ~5 | 6 [avatar split clause 288,168,163,158,244]
295. ~v11_waybel_1(sK0) | v2_struct_0(sK0) | v2_lattice3(sK0) <- (4) [resolution 121,160] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
297. v2_struct_0(sK0) | v2_lattice3(sK0) <- (4, 5) [subsumption resolution 295,165] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
298. v2_lattice3(sK0) <- (4, 5, ~6) [subsumption resolution 297,170] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
299. 18 | ~4 | ~5 | 6 [avatar split clause 298,168,163,158,251]
357. k3_yellow_0(X3) = k12_lattice3(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) | ~r6_waybel_1(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | v2_struct_0(X3) [superposition 135,127] {a:1,w:37,wCS:14800,sine:255,thAx:0,allAx:2,thDist:-2}
362. k3_yellow_0(X3) = k12_lattice3(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) | ~r6_waybel_1(X3,X4,X5) | v2_struct_0(X3) [duplicate literal removal 357] {a:1,w:27,wCS:10800,sine:255,thAx:0,allAx:2,thDist:-2}
363. k3_yellow_0(X3) = k12_lattice3(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v2_lattice3(X3) | ~v5_orders_2(X3) | ~r6_waybel_1(X3,X4,X5) [subsumption resolution 362,101] {a:1,w:25,wCS:10000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
365. k4_yellow_0(X3) = k13_lattice3(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v1_lattice3(X3) | ~v5_orders_2(X3) | ~r6_waybel_1(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | v2_struct_0(X3) [superposition 136,126] {a:1,w:37,wCS:14800,sine:255,thAx:0,allAx:2,thDist:-2}
370. k4_yellow_0(X3) = k13_lattice3(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v1_lattice3(X3) | ~v5_orders_2(X3) | ~r6_waybel_1(X3,X4,X5) | v2_struct_0(X3) [duplicate literal removal 365] {a:1,w:27,wCS:10800,sine:255,thAx:0,allAx:2,thDist:-2}
371. k4_yellow_0(X3) = k13_lattice3(X3,X4,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_orders_2(X3) | ~v1_lattice3(X3) | ~v5_orders_2(X3) | ~r6_waybel_1(X3,X4,X5) [subsumption resolution 370,100] {a:1,w:25,wCS:10000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
403. 30 <=> r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) [avatar definition]
404. r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (30) [avatar component clause 403] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
405. ~r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~30) [avatar component clause 403] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
407. 31 <=> m1_subset_1(k7_waybel_1(sK0,sK1),u1_struct_0(sK0)) [avatar definition]
408. m1_subset_1(k7_waybel_1(sK0,sK1),u1_struct_0(sK0)) <- (31) [avatar component clause 407] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
409. ~m1_subset_1(k7_waybel_1(sK0,sK1),u1_struct_0(sK0)) <- (~31) [avatar component clause 407] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
411. r6_waybel_1(X0,X1,k7_waybel_1(X0,X1)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | v2_struct_0(X0) [resolution 142,132] {a:1,w:32,wCS:12800,sine:1,thAx:0,allAx:2,thDist:-2}
412. r6_waybel_1(X0,X1,k7_waybel_1(X0,X1)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) | v2_struct_0(X0) [duplicate literal removal 411] {a:1,w:26,wCS:10400,sine:1,thAx:0,allAx:2,thDist:-2}
413. r6_waybel_1(X0,X1,k7_waybel_1(X0,X1)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v11_waybel_1(X0) | ~v2_lattice3(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) | ~v4_orders_2(X0) | ~v3_orders_2(X0) [subsumption resolution 412,100] {a:1,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
416. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,X1) | ~r6_waybel_1(sK0,sK1,X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v11_waybel_1(sK0) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1) [superposition 146,129] {a:1,w:33,wCS:3300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
431. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,X1) | ~r6_waybel_1(sK0,sK1,X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v11_waybel_1(sK0) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 3) [subsumption resolution 416,155] {a:1,w:29,wCS:2900,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
432. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,X1) | ~r6_waybel_1(sK0,sK1,X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~v11_waybel_1(sK0) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 3, 4) [subsumption resolution 431,160] {a:1,w:27,wCS:2700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
433. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,X1) | ~r6_waybel_1(sK0,sK1,X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 3, 4, 5) [subsumption resolution 432,165] {a:1,w:25,wCS:2500,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
434. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,X1) | ~r6_waybel_1(sK0,sK1,X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 3, 4, 5, 18) [subsumption resolution 433,253] {a:1,w:23,wCS:2300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
435. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,X1) | ~r6_waybel_1(sK0,sK1,X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 3, 4, 5, 17, 18) [subsumption resolution 434,246] {a:1,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
436. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,X1) | ~r6_waybel_1(sK0,sK1,X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~1, 3, 4, 5, 16, 17, 18) [subsumption resolution 435,239] {a:1,w:19,wCS:1900,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
437. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,X1) | ~r6_waybel_1(sK0,sK1,X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) | ~v3_orders_2(sK0) <- (~1, 3, 4, 5, 15, 16, 17, 18) [subsumption resolution 436,232] {a:1,w:17,wCS:1700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
438. k3_yellow_0(sK0) != k12_lattice3(sK0,sK1,X1) | ~r6_waybel_1(sK0,sK1,X1) | ~m1_subset_1(X1,u1_struct_0(sK0)) <- (~1, 3, 4, 5, 14, 15, 16, 17, 18) [subsumption resolution 437,225] {a:1,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
439. k3_yellow_0(sK0) != k3_yellow_0(sK0) | ~r6_waybel_1(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) | ~r6_waybel_1(sK0,sK1,X0) <- (~1, 3, 4, 5, 14, 15, 16, 17, 18) [superposition 438,363] {a:2,w:31,wCS:3100,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
444. k3_yellow_0(sK0) != k3_yellow_0(sK0) | ~r6_waybel_1(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (~1, 3, 4, 5, 14, 15, 16, 17, 18) [duplicate literal removal 439] {a:2,w:23,wCS:2300,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
445. ~r6_waybel_1(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (~1, 3, 4, 5, 14, 15, 16, 17, 18) [trivial inequality removal 444] {a:2,w:18,wCS:1800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
446. ~r6_waybel_1(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (~1, 3, 4, 5, 14, 15, 16, 17, 18) [subsumption resolution 445,155] {a:2,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
447. ~r6_waybel_1(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v2_lattice3(sK0) | ~v5_orders_2(sK0) <- (~1, 3, 4, 5, 14, 15, 16, 17, 18) [subsumption resolution 446,160] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
448. ~r6_waybel_1(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v5_orders_2(sK0) <- (~1, 3, 4, 5, 14, 15, 16, 17, 18) [subsumption resolution 447,253] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
449. ~r6_waybel_1(sK0,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 3, 4, 5, 14, 15, 16, 17, 18) [subsumption resolution 448,239] {a:2,w:8,wCS:800,nSel:1,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
459. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v11_waybel_1(sK0) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (~30) [resolution 413,405] {a:3,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
464. ~l1_orders_2(sK0) | ~v11_waybel_1(sK0) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (3, ~30) [subsumption resolution 459,155] {a:3,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
465. ~v11_waybel_1(sK0) | ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (3, 4, ~30) [subsumption resolution 464,160] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
466. ~v2_lattice3(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (3, 4, 5, ~30) [subsumption resolution 465,165] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
467. ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (3, 4, 5, 18, ~30) [subsumption resolution 466,253] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
468. ~v5_orders_2(sK0) | ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (3, 4, 5, 17, 18, ~30) [subsumption resolution 467,246] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
469. ~v4_orders_2(sK0) | ~v3_orders_2(sK0) <- (3, 4, 5, 16, 17, 18, ~30) [subsumption resolution 468,239] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
470. ~v3_orders_2(sK0) <- (3, 4, 5, 15, 16, 17, 18, ~30) [subsumption resolution 469,232] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
471. $false <- (3, 4, 5, 14, 15, 16, 17, 18, ~30) [subsumption resolution 470,225] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
472. ~3 | ~4 | ~5 | ~14 | ~15 | ~16 | ~17 | ~18 | 30 [avatar contradiction clause 471]
482. ~m1_subset_1(k7_waybel_1(sK0,sK1),u1_struct_0(sK0)) <- (~1, 3, 4, 5, 14, 15, 16, 17, 18, 30) [resolution 404,449] {a:3,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:11,thDist:-11}
484. ~31 | 1 | ~3 | ~4 | ~5 | ~14 | ~15 | ~16 | ~17 | ~18 | ~30 [avatar split clause 482,403,251,244,237,230,223,163,158,153,144,407]
496. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | v2_struct_0(sK0) <- (~31) [resolution 409,132] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
500. ~l1_orders_2(sK0) | v2_struct_0(sK0) <- (3, ~31) [subsumption resolution 496,155] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
501. v2_struct_0(sK0) <- (3, 4, ~31) [subsumption resolution 500,160] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
502. $false <- (3, 4, ~6, ~31) [subsumption resolution 501,170] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
503. ~3 | ~4 | 6 | 31 [avatar contradiction clause 502]
509. k4_yellow_0(sK0) != k4_yellow_0(sK0) | ~m1_subset_1(k7_waybel_1(sK0,sK1),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~2) [superposition 150,371] {a:2,w:27,wCS:2700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
510. ~m1_subset_1(k7_waybel_1(sK0,sK1),u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~2) [trivial inequality removal 509] {a:2,w:22,wCS:2200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
520. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~2, 31) [subsumption resolution 510,408] {a:2,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
521. ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~2, 3, 31) [subsumption resolution 520,155] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
522. ~v1_lattice3(sK0) | ~v5_orders_2(sK0) | ~r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~2, 3, 4, 31) [subsumption resolution 521,160] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
523. ~v5_orders_2(sK0) | ~r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~2, 3, 4, 17, 31) [subsumption resolution 522,246] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
524. ~r6_waybel_1(sK0,sK1,k7_waybel_1(sK0,sK1)) <- (~2, 3, 4, 16, 17, 31) [subsumption resolution 523,239] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
525. $false <- (~2, 3, 4, 16, 17, 30, 31) [subsumption resolution 524,404] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
526. 2 | ~3 | ~4 | ~16 | ~17 | ~30 | ~31 [avatar contradiction clause 525]
527. $false [avatar sat refutation 151,156,161,166,171,259,269,279,289,299,472,484,503,526] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:379,thDist:-379}
% SZS output end Proof for yellow_5__t34_yellow_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 31
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 51
% Generated clauses: 296
% Active clauses: 68
% Passive clauses: 84
% Final active clauses: 62
% Final passive clauses: 11
% Main loop iterations started: 70
% 
% >>> Simplifying Inferences
% Duplicate literals: 136
% Trivial inequalities: 3
% Fw subsumption resolutions: 103
% 
% >>> Deletion Inferences
% Simple tautologies: 12
% Forward subsumptions: 20
% Backward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 36
% Unit resulting resolution: 3
% Forward superposition: 18
% Backward superposition: 22
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 11
% Split components: 31
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 42
% SAT solver unit clauses: 5
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 9722
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

