i: [1,56,0,5,1,1,0,0] t7_xtuple_0
a: [3,93,0,5,1,56]
s: 93
i: [1,55,0,5,1,1,0,0] t7_xtuple_0
a: [3,98,0,5,1,55]
s: 98
i: [1,57,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,103,0,2,1,57]
s: 103
i: [1,58,0,2,1,0,0,255] fc1_xboole_0
a: [3,108,0,2,1,58]
s: 108
i: [1,80,0,2,1,0,0,2] rc2_xboole_0
a: [3,118,0,2,1,80]
s: 118
i: [1,81,0,2,1,0,0,2] rc1_xboole_0
a: [3,123,0,2,1,81]
s: 123
i: [1,60,0,3,1,0,0,255] fc2_xboole_0
s: 60
i: [1,59,0,3,1,0,0,255] d2_xboole_0
a: [3,113,0,3,1,59]
s: 113
r: 103
i: [1,77,0,4,1,0,0,255] fc3_xboole_0
s: 77
i: [1,61,0,5,2,0,0,255] t6_boole
s: 61
i: [1,70,0,5,2,0,0,1] t7_boole
s: 70
d: [2,128,1,2,1,1,40,70,98]
a: [3,132,1,2,1,128]
s: 132
i: [1,86,0,5,1,0,0,255] d2_tarski
s: 86
i: [1,88,0,5,1,0,0,255] d2_tarski
s: 88
i: [1,79,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 79
i: [1,69,0,7,3,0,0,2] t8_boole
s: 69
d: [2,136,1,5,1,1,40,79,98]
a: [3,142,1,5,1,136]
s: 142
i: [1,78,0,7,1,0,0,255] commutativity_k2_tarski
s: 78
d: [2,137,1,5,1,0,40,79,86]
s: 137
i: [1,83,0,9,2,0,0,1] d4_tarski
s: 83
i: [1,84,0,9,2,0,0,1] d4_tarski
s: 84
d: [2,138,1,5,1,0,40,79,88]
s: 138
i: [1,62,0,10,1,0,0,255] d5_tarski
s: 62
d: [2,163,1,4,1,0,44,77,62]
s: 163
i: [1,82,0,10,3,0,0,1] d4_tarski
s: 82
i: [1,89,0,11,3,0,0,255] d2_tarski
s: 89
d: [2,174,1,7,2,1,40,82,93]
s: 174
i: [1,67,0,14,3,0,0,1] d4_tarski
s: 67
d: [2,181,2,9,3,1,40,174,82]
s: 181
i: [1,66,0,16,3,0,0,1] d4_tarski
s: 66
i: [1,68,0,17,4,0,0,1] d4_tarski
s: 68
d: [2,187,3,8,2,1,40,181,86]
s: 187
i: [1,75,0,17,3,0,0,255] d2_tarski
s: 75
d: [2,199,4,7,1,1,40,187,86]
s: 199
i: [1,76,0,17,3,0,0,255] d2_tarski
s: 76
i: [1,74,0,23,4,0,0,255] d2_tarski
s: 74
d: [2,200,4,7,1,1,40,187,88]
s: 200
d: [2,153,1,6,2,0,40,83,70]
s: 153
d: [2,215,5,5,1,1,44,200,62]
s: 215
d: [2,223,6,0,0,2,40,215,98]
e: 223
d: [2,222,6,0,0,2,52,98,215]
e: 222
f: 222,55,56
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t7_xtuple_0
% SZS output start Proof for xtuple_0__t7_xtuple_0
1. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(k4_tarski(X0,X1),X2) => r2_hidden(X1,k3_tarski(k3_tarski(X2)))) [input t7_xtuple_0]
2. ~! [X0] : ! [X1] : ! [X2] : (r2_hidden(k4_tarski(X0,X1),X2) => r2_hidden(X1,k3_tarski(k3_tarski(X2)))) [negated conjecture 1]
5. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
7. ! [X0] : ! [X1] : (k3_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)))) [input d4_tarski]
8. ! [X0] : ! [X1] : k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [input d5_tarski]
25. ~! [X0,X1,X2] : (r2_hidden(k4_tarski(X0,X1),X2) => r2_hidden(X1,k3_tarski(k3_tarski(X2)))) [flattening 2]
26. ! [X0,X1] : k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [flattening 8]
27. ! [X0,X1] : (k3_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)))) [flattening 7]
30. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 5]
33. ? [X0,X1,X2] : (~r2_hidden(X1,k3_tarski(k3_tarski(X2))) & r2_hidden(k4_tarski(X0,X1),X2)) [ennf transformation 25]
38. ? [X0,X1,X2] : (~r2_hidden(X1,k3_tarski(k3_tarski(X2))) & r2_hidden(k4_tarski(X0,X1),X2)) => (~r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) & r2_hidden(k4_tarski(sK0,sK1),sK2)) [choice axiom]
39. ~r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) & r2_hidden(k4_tarski(sK0,sK1),sK2) [skolemisation 33,38]
40. ! [X0,X1] : ((k3_tarski(X0) = X1 | ? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3))) & (? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)) | ~r2_hidden(X2,X1))) | k3_tarski(X0) != X1)) [nnf transformation 27]
41. ! [X0,X1] : ((k3_tarski(X0) = X1 | ? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(X2,X4)) | r2_hidden(X2,X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (~r2_hidden(X6,X0) | ~r2_hidden(X5,X6))) & (? [X7] : (r2_hidden(X7,X0) & r2_hidden(X5,X7)) | ~r2_hidden(X5,X1))) | k3_tarski(X0) != X1)) [rectify 40]
42. ! [X1,X0] : (? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(X2,X4)) | r2_hidden(X2,X1))) => ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(sK3(X0,X1),X3)) | ~r2_hidden(sK3(X0,X1),X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(sK3(X0,X1),X4)) | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
43. ! [X1,X0] : (? [X4] : (r2_hidden(X4,X0) & r2_hidden(sK3(X0,X1),X4)) => (r2_hidden(sK4(X0,X1),X0) & r2_hidden(sK3(X0,X1),sK4(X0,X1)))) [choice axiom]
44. ! [X5,X0] : (? [X7] : (r2_hidden(X7,X0) & r2_hidden(X5,X7)) => (r2_hidden(sK5(X0,X5),X0) & r2_hidden(X5,sK5(X0,X5)))) [choice axiom]
45. ! [X0,X1] : ((k3_tarski(X0) = X1 | ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(sK3(X0,X1),X3)) | ~r2_hidden(sK3(X0,X1),X1)) & ((r2_hidden(sK4(X0,X1),X0) & r2_hidden(sK3(X0,X1),sK4(X0,X1))) | r2_hidden(sK3(X0,X1),X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (~r2_hidden(X6,X0) | ~r2_hidden(X5,X6))) & ((r2_hidden(sK5(X0,X5),X0) & r2_hidden(X5,sK5(X0,X5))) | ~r2_hidden(X5,X1))) | k3_tarski(X0) != X1)) [skolemisation 41,44,43,42]
46. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 30]
47. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 46]
48. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 47]
49. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK6(X0,X1,X2) != X1 & sK6(X0,X1,X2) != X0) | ~r2_hidden(sK6(X0,X1,X2),X2)) & (sK6(X0,X1,X2) = X1 | sK6(X0,X1,X2) = X0 | r2_hidden(sK6(X0,X1,X2),X2)))) [choice axiom]
50. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK6(X0,X1,X2) != X1 & sK6(X0,X1,X2) != X0) | ~r2_hidden(sK6(X0,X1,X2),X2)) & (sK6(X0,X1,X2) = X1 | sK6(X0,X1,X2) = X0 | r2_hidden(sK6(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 48,49]
55. r2_hidden(k4_tarski(sK0,sK1),sK2) [cnf transformation 39] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
56. ~r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) [cnf transformation 39] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
62. k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [cnf transformation 26] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
65. r2_hidden(X5,X1) | ~r2_hidden(X6,X0) | ~r2_hidden(X5,X6) | k3_tarski(X0) != X1 [cnf transformation 45] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:4,thDist:-4}
72. r2_hidden(X4,X2) | X0 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 50] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
73. r2_hidden(X4,X2) | X1 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 50] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
82. r2_hidden(X5,k3_tarski(X0)) | ~r2_hidden(X6,X0) | ~r2_hidden(X5,X6) [equality resolution 65] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
85. r2_hidden(X4,X2) | k2_tarski(X0,X4) != X2 [equality resolution 73] {a:0,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
86. r2_hidden(X4,k2_tarski(X0,X4)) [equality resolution 85] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. r2_hidden(X4,X2) | k2_tarski(X4,X1) != X2 [equality resolution 72] {a:0,w:8,wCS:3200,sine:255,thAx:0,allAx:2,thDist:-2}
88. r2_hidden(X4,k2_tarski(X4,X1)) [equality resolution 87] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
91. 1 <=> r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) [avatar definition]
93. ~r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) <- (~1) [avatar component clause 91] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
94. ~1 [avatar split clause 56,91]
96. 2 <=> r2_hidden(k4_tarski(sK0,sK1),sK2) [avatar definition]
98. r2_hidden(k4_tarski(sK0,sK1),sK2) <- (2) [avatar component clause 96] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
99. 2 [avatar split clause 55,96]
174. ~r2_hidden(X12,k3_tarski(sK2)) | ~r2_hidden(sK1,X12) <- (~1) [resolution 82,93] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
181. ~r2_hidden(sK1,X0) | ~r2_hidden(X1,sK2) | ~r2_hidden(X0,X1) <- (~1) [resolution 174,82] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
187. ~r2_hidden(k2_tarski(X1,sK1),X0) | ~r2_hidden(X0,sK2) <- (~1) [resolution 181,86] {a:3,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
200. ~r2_hidden(k2_tarski(k2_tarski(X2,sK1),X3),sK2) <- (~1) [resolution 187,88] {a:4,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
215. ~r2_hidden(k4_tarski(X0,sK1),sK2) <- (~1) [superposition 200,62] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
222. $false <- (~1, 2) [unit resulting resolution 98,215] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:17,thDist:-17}
225. 1 | ~2 [avatar contradiction clause 222]
226. $false [avatar sat refutation 94,99,225] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:62,thDist:-62}
% SZS output end Proof for xtuple_0__t7_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 125
% Active clauses: 39
% Passive clauses: 93
% Final active clauses: 38
% Final passive clauses: 54
% Main loop iterations started: 39
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 2
% Forward subsumptions: 16
% 
% >>> Generating Inferences
% Binary resolution: 48
% Unit resulting resolution: 1
% Forward superposition: 22
% Backward superposition: 14
% Self superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 7
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% resolution: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

