i: [1,236,0,2,1,1,0,0] t131_xreal_1
a: [3,375,0,2,1,236]
s: 375
i: [1,235,0,2,1,1,0,0] t131_xreal_1
a: [3,380,0,2,1,235]
s: 380
i: [1,238,0,3,1,1,0,0] t131_xreal_1
a: [3,365,0,3,1,238]
s: 365
i: [1,237,0,3,1,1,0,0] t131_xreal_1
a: [3,370,0,3,1,237]
s: 370
i: [1,239,0,5,1,1,0,0] t131_xreal_1
a: [3,360,0,5,1,239]
s: 360
i: [1,240,0,2,1,0,0,2] fc1_numbers
a: [3,385,0,2,1,240]
s: 385
i: [1,241,0,2,1,0,0,3] spc1_boole
a: [3,390,0,2,1,241]
s: 390
i: [1,242,0,2,1,0,0,2] fc1_xboole_0
a: [3,395,0,2,1,242]
s: 395
i: [1,251,0,2,1,0,0,3] fc6_ordinal1
a: [3,435,0,2,1,251]
s: 435
i: [1,250,0,2,1,0,0,3] fc6_ordinal1
a: [3,440,0,2,1,250]
s: 440
i: [1,252,0,2,1,0,0,2] spc1_numerals
a: [3,460,0,2,1,252]
s: 460
i: [1,324,0,2,1,0,0,2] rc2_xboole_0
a: [3,472,0,2,1,324]
s: 472
i: [1,325,0,2,1,0,0,4] rc4_ordinal1
a: [3,477,0,2,1,325]
s: 477
i: [1,326,0,2,1,0,0,4] rc1_ordinal1
a: [3,482,0,2,1,326]
s: 482
i: [1,327,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,487,0,2,1,327]
s: 487
i: [1,328,0,2,1,0,0,1] rc1_xreal_0
a: [3,492,0,2,1,328]
s: 492
i: [1,329,0,2,1,0,0,2] rc1_xxreal_0
a: [3,497,0,2,1,329]
s: 497
i: [1,330,0,2,1,0,0,2] rc1_xboole_0
a: [3,502,0,2,1,330]
s: 502
i: [1,334,0,2,1,0,0,255] rc2_ordinal1
a: [3,507,0,2,1,334]
s: 507
i: [1,333,0,2,1,0,0,255] rc2_ordinal1
a: [3,512,0,2,1,333]
s: 512
i: [1,332,0,2,1,0,0,255] rc2_ordinal1
a: [3,517,0,2,1,332]
s: 517
i: [1,331,0,2,1,0,0,255] rc2_ordinal1
a: [3,522,0,2,1,331]
s: 522
i: [1,336,0,2,1,0,0,4] rc5_ordinal1
a: [3,527,0,2,1,336]
s: 527
i: [1,335,0,2,1,0,0,4] rc5_ordinal1
a: [3,532,0,2,1,335]
s: 532
i: [1,338,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,537,0,2,1,338]
s: 537
i: [1,337,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,542,0,2,1,337]
s: 542
i: [1,342,0,2,1,0,0,2] rc3_xreal_0
a: [3,547,0,2,1,342]
s: 547
i: [1,341,0,2,1,0,0,2] rc3_xreal_0
a: [3,552,0,2,1,341]
s: 552
i: [1,340,0,2,1,0,0,2] rc3_xreal_0
a: [3,557,0,2,1,340]
s: 557
i: [1,339,0,2,1,0,0,2] rc3_xreal_0
a: [3,562,0,2,1,339]
s: 562
i: [1,346,0,2,1,0,0,2] rc2_xreal_0
a: [3,567,0,2,1,346]
s: 567
i: [1,345,0,2,1,0,0,2] rc2_xreal_0
a: [3,572,0,2,1,345]
s: 572
i: [1,344,0,2,1,0,0,2] rc2_xreal_0
a: [3,577,0,2,1,344]
s: 577
i: [1,343,0,2,1,0,0,2] rc2_xreal_0
a: [3,582,0,2,1,343]
s: 582
i: [1,348,0,2,1,0,0,2] rc3_xxreal_0
a: [3,587,0,2,1,348]
s: 587
i: [1,347,0,2,1,0,0,2] rc3_xxreal_0
a: [3,592,0,2,1,347]
s: 592
i: [1,350,0,2,1,0,0,2] rc2_xxreal_0
a: [3,597,0,2,1,350]
s: 597
i: [1,349,0,2,1,0,0,2] rc2_xxreal_0
a: [3,602,0,2,1,349]
s: 602
i: [1,354,0,2,1,0,0,2] rc4_xreal_0
a: [3,607,0,2,1,354]
s: 607
i: [1,353,0,2,1,0,0,2] rc4_xreal_0
a: [3,612,0,2,1,353]
s: 612
i: [1,352,0,2,1,0,0,2] rc4_xreal_0
a: [3,617,0,2,1,352]
s: 617
i: [1,351,0,2,1,0,0,2] rc4_xreal_0
a: [3,622,0,2,1,351]
s: 622
i: [1,356,0,2,1,0,0,2] rc4_xxreal_0
a: [3,627,0,2,1,356]
s: 627
i: [1,355,0,2,1,0,0,2] rc4_xxreal_0
a: [3,632,0,2,1,355]
s: 632
i: [1,307,0,3,1,0,0,255] reflexivity_r1_tarski
s: 307
i: [1,243,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,400,0,3,1,243]
s: 400
i: [1,244,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,405,0,3,1,244]
s: 405
r: 395
d: [2,634,0,2,1,2,23,395,405]
a: [3,638,0,2,1,634]
s: 638
i: [1,245,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,410,0,3,1,245]
s: 410
i: [1,246,0,3,1,0,0,2] t1_numerals
a: [3,415,0,3,1,246]
s: 415
i: [1,255,0,3,1,0,0,2] spc1_numerals
a: [3,445,0,3,1,255]
s: 445
i: [1,254,0,3,1,0,0,2] spc1_numerals
a: [3,450,0,3,1,254]
s: 450
d: [2,640,0,3,1,2,22,415,405]
a: [3,644,0,3,1,640]
s: 644
d: [2,646,0,3,1,2,22,450,400]
a: [3,650,0,3,1,646]
s: 650
i: [1,257,0,4,2,0,0,4] cc6_ordinal1
s: 257
i: [1,258,0,4,2,0,0,4] cc2_xcmplx_0
s: 258
d: [2,652,1,2,1,1,40,257,477]
a: [3,657,1,2,1,652]
s: 657
i: [1,259,0,4,2,0,0,4] cc2_xreal_0
s: 259
d: [2,653,1,2,1,1,40,257,527]
a: [3,662,1,2,1,653]
s: 662
i: [1,260,0,4,2,0,0,4] cc2_xxreal_0
s: 260
i: [1,261,0,4,2,0,0,255] cc1_ordinal1
s: 261
d: [2,664,1,2,1,1,40,258,477]
a: [3,669,1,2,1,664]
s: 669
i: [1,262,0,4,2,0,0,255] cc1_ordinal1
s: 262
d: [2,665,1,2,1,1,40,258,527]
a: [3,674,1,2,1,665]
s: 674
i: [1,266,0,4,2,0,0,2] cc3_xreal_0
s: 266
i: [1,267,0,4,2,0,0,2] cc4_xreal_0
s: 267
d: [2,750,1,2,1,1,40,266,380]
a: [3,759,1,2,1,750]
s: 759
i: [1,276,0,4,2,0,0,255] cc9_ordinal1
s: 276
d: [2,751,1,2,1,1,40,266,375]
a: [3,764,1,2,1,751]
s: 764
i: [1,277,0,4,2,0,0,4] cc7_ordinal1
s: 277
i: [1,278,0,4,2,0,0,4] cc3_ordinal1
s: 278
d: [2,771,1,2,1,1,40,267,380]
a: [3,780,1,2,1,771]
s: 780
i: [1,306,0,4,1,0,0,3] existence_m1_subset_1
s: 306
d: [2,772,1,2,1,1,40,267,375]
a: [3,785,1,2,1,772]
s: 785
i: [1,247,0,4,1,0,0,2] dt_k5_numbers
a: [3,420,0,4,1,247]
s: 420
i: [1,248,0,4,1,0,0,1] dt_k6_numbers
a: [3,425,0,4,1,248]
s: 425
d: [2,676,1,2,1,1,40,259,477]
a: [3,681,1,2,1,676]
s: 681
i: [1,253,0,4,1,0,0,2] spc1_numerals
a: [3,455,0,4,1,253]
s: 455
d: [2,677,1,2,1,1,40,259,527]
a: [3,686,1,2,1,677]
s: 686
d: [2,864,0,4,1,2,22,420,400]
a: [3,868,0,4,1,864]
s: 868
d: [2,870,0,4,1,2,22,425,400]
a: [3,874,0,4,1,870]
s: 874
d: [2,688,1,2,1,1,40,260,477]
a: [3,693,1,2,1,688]
s: 693
d: [2,879,0,4,1,2,22,455,400]
a: [3,883,0,4,1,879]
s: 883
d: [2,689,1,2,1,1,40,260,527]
a: [3,698,1,2,1,689]
s: 698
i: [1,279,0,5,2,0,0,2] t6_boole
s: 279
i: [1,280,0,5,2,0,0,3] cc8_ordinal1
s: 280
d: [2,700,1,2,1,1,40,261,435]
a: [3,708,1,2,1,700]
s: 708
i: [1,281,0,5,2,0,0,2] cc1_xcmplx_0
s: 281
d: [2,702,1,2,1,1,40,261,482]
a: [3,718,1,2,1,702]
s: 718
i: [1,282,0,5,2,0,0,2] cc1_xreal_0
s: 282
i: [1,303,0,5,2,0,0,255] t7_boole
s: 303
d: [2,725,1,2,1,1,40,262,435]
a: [3,733,1,2,1,725]
s: 733
i: [1,249,0,5,1,0,0,3] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,430,0,5,1,249]
s: 430
d: [2,727,1,2,1,1,40,262,482]
a: [3,743,1,2,1,727]
s: 743
i: [1,321,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,467,0,5,2,321]
s: 467
d: [2,888,0,5,2,1,22,279,405]
s: 888
d: [2,929,1,5,2,2,44,365,888]
s: 929
i: [1,287,0,6,3,0,0,255] cc2_ordinal1
s: 287
d: [2,930,1,5,2,2,44,370,888]
s: 930
i: [1,290,0,6,3,0,0,2] cc5_xxreal_0
s: 290
i: [1,292,0,6,3,0,0,2] cc5_xxreal_0
s: 292
d: [2,928,1,7,2,2,44,360,888]
s: 928
i: [1,293,0,6,3,0,0,2] cc3_xxreal_0
s: 293
d: [2,975,2,8,2,3,40,928,467]
a: [3,984,2,4,1,975]
s: 984
i: [1,295,0,6,3,0,0,2] cc3_xxreal_0
s: 295
i: [1,297,0,6,3,0,0,2] cc7_xxreal_0
s: 297
d: [2,752,1,2,1,1,40,266,492]
a: [3,769,1,2,1,752]
s: 769
i: [1,298,0,6,3,0,0,2] cc7_xxreal_0
s: 298
d: [2,773,1,2,1,1,40,267,492]
a: [3,790,1,2,1,773]
s: 790
i: [1,299,0,6,2,0,0,255] t1_subset
s: 299
i: [1,308,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 308
d: [2,792,1,2,1,1,40,276,638]
a: [3,799,1,2,1,792]
s: 799
i: [1,263,0,7,3,0,0,4] cc5_ordinal1
s: 263
d: [2,793,1,2,1,1,40,276,502]
a: [3,804,1,2,1,793]
s: 804
i: [1,264,0,7,2,0,0,1] t2_arithm
s: 264
i: [1,265,0,7,2,0,0,3] t3_arithm
s: 265
d: [2,794,1,2,1,1,40,276,622]
a: [3,809,1,2,1,794]
s: 809
i: [1,301,0,7,2,0,0,255] t3_subset
s: 301
d: [2,795,1,2,1,1,40,276,632]
a: [3,814,1,2,1,795]
s: 814
i: [1,302,0,7,3,0,0,2] t8_boole
s: 302
i: [1,256,0,8,3,0,0,255] cc10_ordinal1
s: 256
d: [2,816,1,2,1,1,40,277,638]
a: [3,823,1,2,1,816]
s: 823
i: [1,284,0,8,4,0,0,2] cc4_xxreal_0
s: 284
d: [2,976,2,6,2,3,40,928,930]
a: [3,988,2,2,1,976]
s: 988
r: 988
i: [1,286,0,8,4,0,0,2] cc6_xxreal_0
s: 286
i: [1,288,0,8,4,0,0,2] cc8_xxreal_0
s: 288
a: [3,989,2,2,1,976]
s: 989
i: [1,300,0,8,3,0,0,255] t2_subset
s: 300
d: [2,1060,2,6,3,1,40,284,780]
a: [3,1075,2,2,1,1060]
s: 1075
i: [1,318,0,8,3,0,0,1] fc3_xcmplx_0
s: 318
i: [1,320,0,8,3,0,0,1] fc6_xreal_0
s: 320
a: [3,1078,2,2,1,1060]
s: 1078
i: [1,305,0,9,3,0,0,255] t5_subset
s: 305
d: [2,1061,2,6,3,1,40,284,785]
a: [3,1084,2,2,1,1061]
s: 1084
r: 1084
i: [1,304,0,10,3,0,0,255] t4_subset
s: 304
i: [1,322,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 322
a: [3,1083,2,2,1,1061]
s: 1083
i: [1,270,0,11,5,0,0,1] t2_real
s: 270
a: [3,1088,2,2,1,1061]
s: 1088
r: 1088
i: [1,271,0,11,5,0,0,1] t1_real
s: 271
i: [1,272,0,11,5,0,0,1] t8_real
s: 272
a: [3,1087,2,2,1,1061]
s: 1087
i: [1,273,0,11,5,0,0,1] t7_real
s: 273
a: [3,1092,2,2,1,1061]
s: 1092
i: [1,274,0,11,5,0,0,1] t3_real
s: 274
i: [1,275,0,11,5,0,0,1] t4_real
s: 275
d: [2,1295,1,12,4,1,40,272,360]
a: [3,1304,1,4,1,1295]
s: 1304
d: [2,1315,2,0,0,3,52,380,375,1304,320]
e: 1315
d: [2,1316,2,4,2,1,40,1304,320]
d: [2,1318,2,2,1,2,21,1316,375]
d: [2,1319,2,0,0,3,21,1318,380]
e: 1319
r: 1304
i: [1,319,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 319
a: [3,1303,1,4,1,1295]
s: 1303
i: [1,309,0,12,5,0,0,1] fc26_xreal_0
s: 309
i: [1,310,0,12,5,0,0,1] fc24_xreal_0
s: 310
a: [3,1308,1,4,1,1295]
s: 1308
d: [2,1381,3,0,0,5,52,1087,375,1078,380,1308,310]
e: 1381
d: [2,1382,2,8,4,1,40,1308,310]
d: [2,1386,2,6,3,2,21,1382,380]
d: [2,1387,2,4,2,3,21,1386,1078]
d: [2,1388,2,2,1,4,21,1387,375]
d: [2,1389,2,0,0,5,21,1388,1087]
e: 1389
r: 1075
r: 1078
r: 1087
r: 1092
i: [1,311,0,12,5,0,0,1] fc23_xreal_0
s: 311
r: 1308
d: [2,1252,1,8,4,1,40,270,365]
d: [2,1258,1,6,3,2,21,1252,375]
a: [3,1261,1,2,1,1258]
s: 1261
r: 1261
i: [1,312,0,12,5,0,0,1] fc25_xreal_0
s: 312
i: [1,317,0,12,5,0,0,1] fc8_xcmplx_0
s: 317
a: [3,1262,1,2,1,1258]
s: 1262
i: [1,268,0,13,6,0,0,1] t5_real
s: 268
d: [2,1278,1,8,4,1,40,271,365]
d: [2,1284,1,6,3,2,21,1278,375]
a: [3,1288,1,2,1,1284]
s: 1288
d: [2,889,1,2,1,1,40,280,644]
d: [2,1039,2,2,1,1,40,823,260]
d: [2,1040,2,2,1,1,40,823,259]
d: [2,1250,1,8,4,1,40,270,370]
d: [2,1401,1,6,3,2,21,1250,380]
d: [2,1398,2,4,2,3,21,1386,375]
d: [2,1426,2,4,2,1,40,1261,290]
d: [2,1428,2,2,1,2,21,1426,638]
d: [2,1474,2,4,2,1,40,1288,293]
d: [2,1477,2,2,1,2,21,1474,638]
f: 1477,1428,1398,1401,1319,1295,1284,1040,1039,889,640,634,246,244,242,235,236,237,238,239
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t131_xreal_1
% SZS output start Proof for xreal_1__t131_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0)) => r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers)))) [input t131_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0)) => r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers)))) [negated conjecture 1]
12. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
13. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
16. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
18. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
21. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v2_xxreal_0(k3_xcmplx_0(X1,X0))) [input fc24_xreal_0]
29. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
31. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
39. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
52. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
59. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
74. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
85. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
90. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k3_xcmplx_0(X0,X1))) [input fc6_xreal_0]
107. ? [X0] : (? [X1] : ((~r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & (r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
108. ? [X0] : (? [X1] : (~r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 107]
112. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 12]
113. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 13]
124. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 31]
125. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 124]
126. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 74]
127. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 126]
128. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 39]
129. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 128]
140. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 18]
151. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 16]
152. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 151]
153. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 52]
154. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 153]
169. ! [X0,X1] : (~v2_xxreal_0(k3_xcmplx_0(X1,X0)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 21]
170. ! [X0,X1] : (~v2_xxreal_0(k3_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 169]
187. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 90]
188. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 187]
195. ? [X0] : (? [X1] : (~r1_xxreal_0(k3_xcmplx_0(X0,X1),k6_numbers) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~r1_xxreal_0(k3_xcmplx_0(sK0,X1),k6_numbers) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
196. ? [X1] : (~r1_xxreal_0(k3_xcmplx_0(sK0,X1),k6_numbers) & r1_xxreal_0(X1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) => (~r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) & r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) [choice axiom]
197. (~r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) & r1_xxreal_0(sK1,k6_numbers) & r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 108,196,195]
235. v1_xreal_0(sK0) [cnf transformation 197] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
236. v1_xreal_0(sK1) [cnf transformation 197] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
237. r1_xxreal_0(k6_numbers,sK0) [cnf transformation 197] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
238. r1_xxreal_0(sK1,k6_numbers) [cnf transformation 197] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
239. ~r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) [cnf transformation 197] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
242. v1_xboole_0(k1_xboole_0) [cnf transformation 59] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
244. k6_numbers = k1_xboole_0 [cnf transformation 85] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
246. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 29] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
259. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 112] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
260. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 113] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
270. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 125] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
271. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 127] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
272. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 129] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
280. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 140] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
290. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 152] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
293. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 154] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
310. ~v2_xxreal_0(k3_xcmplx_0(X1,X0)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 170] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
320. v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 188] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
358. 1 <=> r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) [avatar definition]
360. ~r1_xxreal_0(k3_xcmplx_0(sK0,sK1),k6_numbers) <- (~1) [avatar component clause 358] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
361. ~1 [avatar split clause 239,358]
363. 2 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
365. r1_xxreal_0(sK1,k6_numbers) <- (2) [avatar component clause 363] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
366. 2 [avatar split clause 238,363]
368. 3 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
370. r1_xxreal_0(k6_numbers,sK0) <- (3) [avatar component clause 368] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
371. 3 [avatar split clause 237,368]
373. 4 <=> v1_xreal_0(sK1) [avatar definition]
375. v1_xreal_0(sK1) <- (4) [avatar component clause 373] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
376. 4 [avatar split clause 236,373]
378. 5 <=> v1_xreal_0(sK0) [avatar definition]
380. v1_xreal_0(sK0) <- (5) [avatar component clause 378] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
381. 5 [avatar split clause 235,378]
393. 8 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
395. v1_xboole_0(k1_xboole_0) <- (8) [avatar component clause 393] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
396. 8 [avatar split clause 242,393]
403. 10 <=> k6_numbers = k1_xboole_0 [avatar definition]
405. k6_numbers = k1_xboole_0 <- (10) [avatar component clause 403] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
406. 10 [avatar split clause 244,403]
413. 12 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
415. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (12) [avatar component clause 413] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
416. 12 [avatar split clause 246,413]
634. v1_xboole_0(k6_numbers) <- (8, 10) [backward demodulation 395,405] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
636. 57 <=> v1_xboole_0(k6_numbers) [avatar definition]
638. v1_xboole_0(k6_numbers) <- (57) [avatar component clause 636] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
639. 57 | ~8 | ~10 [avatar split clause 634,403,393,636]
640. m1_subset_1(k6_numbers,k4_ordinal1) <- (10, 12) [forward demodulation 415,405] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
642. 58 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
644. m1_subset_1(k6_numbers,k4_ordinal1) <- (58) [avatar component clause 642] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
645. 58 | ~10 | ~12 [avatar split clause 640,413,403,642]
821. 86 <=> v7_ordinal1(k6_numbers) [avatar definition]
823. v7_ordinal1(k6_numbers) <- (86) [avatar component clause 821] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
889. v7_ordinal1(k6_numbers) <- (58) [resolution 280,644] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
892. 86 | ~58 [avatar split clause 889,642,821]
1039. v1_xxreal_0(k6_numbers) <- (86) [resolution 823,260] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1040. v1_xreal_0(k6_numbers) <- (86) [resolution 823,259] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1045. 118 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1048. 118 | ~86 [avatar split clause 1039,821,1045]
1050. 119 <=> v1_xreal_0(k6_numbers) [avatar definition]
1053. 119 | ~86 [avatar split clause 1040,821,1050]
1077. 122 <=> v3_xxreal_0(sK0) [avatar definition]
1086. 124 <=> v2_xxreal_0(sK1) [avatar definition]
1250. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (3) [resolution 270,370] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1260. 150 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1261. v3_xxreal_0(k6_numbers) <- (150) [avatar component clause 1260] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1278. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (2) [resolution 271,365] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1284. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 4) [subsumption resolution 1278,375] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1286. 152 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1288. v2_xxreal_0(k6_numbers) <- (152) [avatar component clause 1286] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1289. ~119 | 152 | ~124 | ~2 | ~4 [avatar split clause 1284,373,363,1086,1286,1050]
1295. v3_xxreal_0(k6_numbers) | v2_xxreal_0(k3_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k3_xcmplx_0(sK0,sK1)) <- (~1) [resolution 272,360] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1302. 154 <=> v1_xreal_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
1304. ~v1_xreal_0(k3_xcmplx_0(sK0,sK1)) <- (~154) [avatar component clause 1302] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1306. 155 <=> v2_xxreal_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
1308. v2_xxreal_0(k3_xcmplx_0(sK0,sK1)) <- (155) [avatar component clause 1306] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1309. ~154 | ~119 | 155 | 150 | 1 [avatar split clause 1295,358,1260,1306,1050,1302]
1316. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~154) [resolution 1304,320] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1318. ~v1_xreal_0(sK0) <- (4, ~154) [subsumption resolution 1316,375] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1319. $false <- (4, 5, ~154) [subsumption resolution 1318,380] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1320. ~4 | ~5 | 154 [avatar contradiction clause 1319]
1382. ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) | ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) <- (155) [resolution 1308,310] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1386. v3_xxreal_0(sK0) | ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) <- (5, 155) [subsumption resolution 1382,380] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1398. v3_xxreal_0(sK0) | v2_xxreal_0(sK1) <- (4, 5, 155) [subsumption resolution 1386,375] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1401. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 5) [subsumption resolution 1250,380] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1402. ~119 | 150 | ~122 | ~3 | ~5 [avatar split clause 1401,378,368,1077,1260,1050]
1410. 124 | 122 | ~4 | ~5 | ~155 [avatar split clause 1398,1306,378,373,1077,1086]
1426. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (150) [resolution 1261,290] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1428. ~v1_xxreal_0(k6_numbers) <- (57, 150) [subsumption resolution 1426,638] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1429. ~118 | ~57 | ~150 [avatar split clause 1428,1260,636,1045]
1474. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (152) [resolution 1288,293] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1477. ~v1_xxreal_0(k6_numbers) <- (57, 152) [subsumption resolution 1474,638] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1478. ~118 | ~57 | ~152 [avatar split clause 1477,1286,636,1045]
1479. $false [avatar sat refutation 361,366,371,376,381,396,406,416,639,645,892,1048,1053,1289,1309,1320,1402,1410,1429,1478] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:930,thDist:-930}
% SZS output end Proof for xreal_1__t131_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 94
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 122
% Generated clauses: 595
% Active clauses: 145
% Passive clauses: 253
% Final active clauses: 134
% Final passive clauses: 83
% Main loop iterations started: 159
% 
% >>> Simplifying Inferences
% Duplicate literals: 43
% Fw subsumption resolutions: 49
% Fw demodulations: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 30
% Forward subsumptions: 61
% 
% >>> Generating Inferences
% Binary resolution: 154
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 41
% Backward superposition: 14
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 41
% Split components: 99
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 178
% SAT solver unit clauses: 57
% SAT solver binary clauses: 62
% 
% Memory used [KB]: 5756
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.006 s ( own 0.004 s ) 

