i: [1,224,0,2,1,1,0,0] t140_xreal_1
a: [3,353,0,2,1,224]
s: 353
i: [1,223,0,2,1,1,0,0] t140_xreal_1
a: [3,358,0,2,1,223]
s: 358
i: [1,226,0,3,1,1,0,0] t140_xreal_1
a: [3,343,0,3,1,226]
s: 343
i: [1,225,0,3,1,1,0,0] t140_xreal_1
a: [3,348,0,3,1,225]
s: 348
i: [1,227,0,5,1,1,0,0] t140_xreal_1
a: [3,338,0,5,1,227]
s: 338
i: [1,228,0,2,1,0,0,255] fc1_numbers
a: [3,363,0,2,1,228]
s: 363
i: [1,229,0,2,1,0,0,2] fc1_xboole_0
a: [3,368,0,2,1,229]
s: 368
i: [1,236,0,2,1,0,0,3] fc6_ordinal1
a: [3,398,0,2,1,236]
s: 398
i: [1,235,0,2,1,0,0,3] fc6_ordinal1
a: [3,403,0,2,1,235]
s: 403
i: [1,302,0,2,1,0,0,2] rc2_xboole_0
a: [3,415,0,2,1,302]
s: 415
i: [1,303,0,2,1,0,0,4] rc4_ordinal1
a: [3,420,0,2,1,303]
s: 420
i: [1,304,0,2,1,0,0,4] rc1_ordinal1
a: [3,425,0,2,1,304]
s: 425
i: [1,305,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,430,0,2,1,305]
s: 430
i: [1,306,0,2,1,0,0,1] rc1_xreal_0
a: [3,435,0,2,1,306]
s: 435
i: [1,307,0,2,1,0,0,2] rc1_xxreal_0
a: [3,440,0,2,1,307]
s: 440
i: [1,308,0,2,1,0,0,2] rc1_xboole_0
a: [3,445,0,2,1,308]
s: 445
i: [1,312,0,2,1,0,0,255] rc2_ordinal1
a: [3,450,0,2,1,312]
s: 450
i: [1,311,0,2,1,0,0,255] rc2_ordinal1
a: [3,455,0,2,1,311]
s: 455
i: [1,310,0,2,1,0,0,255] rc2_ordinal1
a: [3,460,0,2,1,310]
s: 460
i: [1,309,0,2,1,0,0,255] rc2_ordinal1
a: [3,465,0,2,1,309]
s: 465
i: [1,314,0,2,1,0,0,4] rc5_ordinal1
a: [3,470,0,2,1,314]
s: 470
i: [1,313,0,2,1,0,0,4] rc5_ordinal1
a: [3,475,0,2,1,313]
s: 475
i: [1,316,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,480,0,2,1,316]
s: 480
i: [1,315,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,485,0,2,1,315]
s: 485
i: [1,320,0,2,1,0,0,2] rc2_xreal_0
a: [3,490,0,2,1,320]
s: 490
i: [1,319,0,2,1,0,0,2] rc2_xreal_0
a: [3,495,0,2,1,319]
s: 495
i: [1,318,0,2,1,0,0,2] rc2_xreal_0
a: [3,500,0,2,1,318]
s: 500
i: [1,317,0,2,1,0,0,2] rc2_xreal_0
a: [3,505,0,2,1,317]
s: 505
i: [1,324,0,2,1,0,0,2] rc3_xreal_0
a: [3,510,0,2,1,324]
s: 510
i: [1,323,0,2,1,0,0,2] rc3_xreal_0
a: [3,515,0,2,1,323]
s: 515
i: [1,322,0,2,1,0,0,2] rc3_xreal_0
a: [3,520,0,2,1,322]
s: 520
i: [1,321,0,2,1,0,0,2] rc3_xreal_0
a: [3,525,0,2,1,321]
s: 525
i: [1,326,0,2,1,0,0,2] rc2_xxreal_0
a: [3,530,0,2,1,326]
s: 530
i: [1,325,0,2,1,0,0,2] rc2_xxreal_0
a: [3,535,0,2,1,325]
s: 535
i: [1,328,0,2,1,0,0,2] rc3_xxreal_0
a: [3,540,0,2,1,328]
s: 540
i: [1,327,0,2,1,0,0,2] rc3_xxreal_0
a: [3,545,0,2,1,327]
s: 545
i: [1,332,0,2,1,0,0,2] rc4_xreal_0
a: [3,550,0,2,1,332]
s: 550
i: [1,331,0,2,1,0,0,2] rc4_xreal_0
a: [3,555,0,2,1,331]
s: 555
i: [1,330,0,2,1,0,0,2] rc4_xreal_0
a: [3,560,0,2,1,330]
s: 560
i: [1,329,0,2,1,0,0,2] rc4_xreal_0
a: [3,565,0,2,1,329]
s: 565
i: [1,334,0,2,1,0,0,2] rc4_xxreal_0
a: [3,570,0,2,1,334]
s: 570
i: [1,333,0,2,1,0,0,2] rc4_xxreal_0
a: [3,575,0,2,1,333]
s: 575
i: [1,287,0,3,1,0,0,255] reflexivity_r1_tarski
s: 287
i: [1,230,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,373,0,3,1,230]
s: 373
i: [1,231,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,378,0,3,1,231]
s: 378
r: 368
d: [2,577,0,2,1,2,23,368,378]
a: [3,581,0,2,1,577]
s: 581
i: [1,232,0,3,1,0,0,2] t1_numerals
a: [3,383,0,3,1,232]
s: 383
d: [2,583,0,3,1,2,22,383,378]
a: [3,587,0,3,1,583]
s: 587
i: [1,238,0,4,2,0,0,4] cc6_ordinal1
s: 238
i: [1,239,0,4,2,0,0,4] cc2_xcmplx_0
s: 239
i: [1,240,0,4,2,0,0,4] cc2_xreal_0
s: 240
d: [2,589,1,2,1,1,40,238,420]
a: [3,594,1,2,1,589]
s: 594
i: [1,241,0,4,2,0,0,4] cc2_xxreal_0
s: 241
d: [2,590,1,2,1,1,40,238,470]
a: [3,599,1,2,1,590]
s: 599
i: [1,242,0,4,2,0,0,255] cc1_ordinal1
s: 242
i: [1,243,0,4,2,0,0,255] cc1_ordinal1
s: 243
d: [2,601,1,2,1,1,40,239,420]
a: [3,606,1,2,1,601]
s: 606
i: [1,246,0,4,2,0,0,2] cc3_xreal_0
s: 246
d: [2,687,1,2,1,1,40,246,358]
a: [3,696,1,2,1,687]
s: 696
i: [1,247,0,4,2,0,0,2] cc4_xreal_0
s: 247
i: [1,256,0,4,2,0,0,255] cc9_ordinal1
s: 256
d: [2,688,1,2,1,1,40,246,353]
a: [3,701,1,2,1,688]
s: 701
i: [1,257,0,4,2,0,0,4] cc7_ordinal1
s: 257
d: [2,708,1,2,1,1,40,247,358]
a: [3,717,1,2,1,708]
s: 717
i: [1,258,0,4,2,0,0,4] cc3_ordinal1
s: 258
i: [1,286,0,4,1,0,0,3] existence_m1_subset_1
s: 286
d: [2,709,1,2,1,1,40,247,353]
a: [3,722,1,2,1,709]
s: 722
i: [1,233,0,4,1,0,0,255] dt_k5_numbers
a: [3,388,0,4,1,233]
s: 388
d: [2,602,1,2,1,1,40,239,470]
a: [3,611,1,2,1,602]
s: 611
i: [1,234,0,4,1,0,0,255] dt_k6_numbers
a: [3,393,0,4,1,234]
s: 393
d: [2,801,0,4,1,2,22,388,373]
a: [3,805,0,4,1,801]
s: 805
d: [2,613,1,2,1,1,40,240,420]
a: [3,618,1,2,1,613]
s: 618
d: [2,807,0,4,1,2,22,393,373]
a: [3,811,0,4,1,807]
s: 811
d: [2,614,1,2,1,1,40,240,470]
a: [3,623,1,2,1,614]
s: 623
i: [1,259,0,5,2,0,0,2] t6_boole
s: 259
i: [1,260,0,5,2,0,0,3] cc8_ordinal1
s: 260
d: [2,625,1,2,1,1,40,241,420]
a: [3,630,1,2,1,625]
s: 630
i: [1,261,0,5,2,0,0,255] cc1_xcmplx_0
s: 261
d: [2,626,1,2,1,1,40,241,470]
a: [3,635,1,2,1,626]
s: 635
i: [1,262,0,5,2,0,0,255] cc1_xreal_0
s: 262
i: [1,283,0,5,2,0,0,255] t7_boole
s: 283
d: [2,637,1,2,1,1,40,242,398]
a: [3,645,1,2,1,637]
s: 645
i: [1,300,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,410,0,5,2,300]
s: 410
d: [2,639,1,2,1,1,40,242,425]
a: [3,655,1,2,1,639]
s: 655
d: [2,819,0,5,2,1,22,259,378]
s: 819
i: [1,267,0,6,3,0,0,255] cc2_ordinal1
s: 267
d: [2,842,1,5,2,2,44,343,819]
s: 842
i: [1,270,0,6,3,0,0,2] cc5_xxreal_0
s: 270
d: [2,848,2,4,2,3,40,842,410]
d: [2,849,2,2,1,4,21,848,722]
a: [3,853,2,2,1,849]
s: 853
i: [1,272,0,6,3,0,0,2] cc5_xxreal_0
s: 272
i: [1,273,0,6,3,0,0,2] cc3_xxreal_0
s: 273
d: [2,843,1,5,2,2,44,348,819]
s: 843
i: [1,275,0,6,3,0,0,2] cc3_xxreal_0
s: 275
d: [2,897,2,4,2,3,40,843,410]
d: [2,898,2,2,1,4,21,897,717]
a: [3,902,2,2,1,898]
s: 902
i: [1,277,0,6,3,0,0,2] cc7_xxreal_0
s: 277
i: [1,278,0,6,3,0,0,2] cc7_xxreal_0
s: 278
d: [2,841,1,7,2,2,44,338,819]
s: 841
i: [1,279,0,6,2,0,0,255] t1_subset
s: 279
d: [2,662,1,2,1,1,40,243,398]
a: [3,670,1,2,1,662]
s: 670
i: [1,288,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 288
i: [1,244,0,7,3,0,0,4] cc5_ordinal1
s: 244
d: [2,664,1,2,1,1,40,243,425]
a: [3,680,1,2,1,664]
s: 680
i: [1,245,0,7,2,0,0,1] t5_arithm
s: 245
d: [2,689,1,2,1,1,40,246,435]
a: [3,706,1,2,1,689]
s: 706
i: [1,281,0,7,2,0,0,255] t3_subset
s: 281
i: [1,282,0,7,3,0,0,2] t8_boole
s: 282
d: [2,710,1,2,1,1,40,247,435]
a: [3,727,1,2,1,710]
s: 727
i: [1,237,0,8,3,0,0,255] cc10_ordinal1
s: 237
d: [2,729,1,2,1,1,40,256,581]
a: [3,736,1,2,1,729]
s: 736
i: [1,264,0,8,4,0,0,2] cc4_xxreal_0
s: 264
r: 410
r: 853
r: 902
a: [3,407,0,2,1,300]
s: 407
r: 270
r: 272
r: 273
r: 264
d: [2,1031,1,0,0,2,40,407,570]
e: 1031
d: [2,1030,1,0,0,2,40,407,555]
e: 1030
d: [2,1029,1,0,0,2,40,407,545]
e: 1029
d: [2,1028,1,0,0,2,40,407,535]
e: 1028
d: [2,1027,1,0,0,2,40,407,520]
e: 1027
d: [2,1026,1,0,0,2,40,407,500]
e: 1026
d: [2,1025,2,0,0,2,40,407,635]
e: 1025
d: [2,1024,1,0,0,2,40,407,440]
e: 1024
d: [2,1023,2,0,0,2,40,407,727]
e: 1023
d: [2,1022,2,0,0,2,40,407,630]
e: 1022
d: [2,1021,2,0,0,2,40,407,722]
e: 1021
d: [2,1020,2,0,0,2,40,407,717]
e: 1020
d: [2,1017,2,0,0,2,52,717,407]
e: 1017
d: [2,1018,1,0,0,2,52,358,407,247]
e: 1018
d: [2,1019,1,0,0,2,52,420,407,241]
e: 1019
r: 407
s: 853
s: 410
s: 902
s: 270
s: 272
d: [2,928,2,6,3,1,40,264,717]
d: [2,940,2,4,2,2,21,928,902]
a: [3,944,2,2,1,940]
s: 944
s: 273
d: [2,929,2,6,3,1,40,264,722]
d: [2,950,2,4,2,2,21,929,853]
a: [3,954,2,2,1,950]
s: 954
s: 264
i: [1,266,0,8,4,0,0,2] cc6_xxreal_0
s: 266
d: [2,904,2,6,2,3,40,841,842]
a: [3,1013,2,4,1,904]
s: 1013
i: [1,268,0,8,4,0,0,2] cc8_xxreal_0
s: 268
d: [2,730,1,2,1,1,40,256,445]
a: [3,741,1,2,1,730]
s: 741
i: [1,280,0,8,3,0,0,255] t2_subset
s: 280
i: [1,298,0,8,3,0,0,1] fc5_xcmplx_0
s: 298
d: [2,731,1,2,1,1,40,256,565]
a: [3,746,1,2,1,731]
s: 746
i: [1,299,0,8,3,0,0,1] fc8_xreal_0
s: 299
d: [2,732,1,2,1,1,40,256,575]
a: [3,751,1,2,1,732]
s: 751
i: [1,285,0,9,3,0,0,255] t5_subset
s: 285
i: [1,284,0,10,3,0,0,255] t4_subset
s: 284
d: [2,753,1,2,1,1,40,257,581]
a: [3,760,1,2,1,753]
s: 760
i: [1,301,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 301
d: [2,1139,1,7,3,1,40,301,343]
d: [2,1145,1,5,2,2,21,1139,722]
a: [3,1149,1,3,1,1145]
s: 1149
i: [1,250,0,11,5,0,0,1] t2_real
s: 250
i: [1,251,0,11,5,0,0,1] t1_real
s: 251
r: 954
d: [2,1160,1,12,4,1,40,250,338]
a: [3,1177,1,2,1,1160]
s: 1177
i: [1,252,0,11,5,0,0,1] t8_real
s: 252
d: [2,1187,1,12,4,1,40,251,338]
a: [3,1196,1,2,1,1187]
s: 1196
r: 1196
i: [1,253,0,11,5,0,0,1] t7_real
s: 253
i: [1,254,0,11,5,0,0,1] t4_real
s: 254
a: [3,1195,1,2,1,1187]
s: 1195
i: [1,255,0,11,5,0,0,1] t3_real
s: 255
a: [3,948,2,2,1,940]
s: 948
d: [2,1237,3,0,0,3,52,717,944,948,272]
e: 1237
d: [2,1238,3,4,2,1,40,948,272]
d: [2,1241,3,2,1,2,21,1238,944]
d: [2,1242,3,0,0,3,21,1241,717]
e: 1242
r: 944
i: [1,289,0,12,5,0,0,1] fc31_xreal_0
s: 289
i: [1,290,0,12,5,0,0,1] fc32_xreal_0
s: 290
a: [3,943,2,2,1,940]
s: 943
i: [1,291,0,12,5,0,0,1] fc34_xreal_0
s: 291
a: [3,953,2,2,1,950]
s: 953
i: [1,292,0,12,5,0,0,1] fc33_xreal_0
s: 292
i: [1,297,0,12,5,0,0,1] fc9_xcmplx_0
s: 297
a: [3,958,2,2,1,950]
s: 958
i: [1,248,0,13,6,0,0,1] t6_real
s: 248
d: [2,1140,1,7,3,1,40,301,348]
d: [2,1151,1,5,2,2,21,1140,717]
a: [3,1155,1,3,1,1151]
s: 1155
i: [1,249,0,13,6,0,0,1] t5_real
s: 249
i: [1,296,0,14,4,0,0,255] existence_m2_subset_1
s: 296
a: [3,1169,1,4,1,1160]
s: 1169
d: [2,1288,2,0,0,3,52,358,353,1169,299]
e: 1288
d: [2,1289,2,4,2,1,40,1169,299]
d: [2,1291,2,2,1,2,21,1289,353]
d: [2,1292,2,0,0,3,21,1291,358]
e: 1292
r: 1169
i: [1,293,0,15,5,0,0,255] dt_m2_subset_1
s: 293
a: [3,1168,1,4,1,1160]
s: 1168
i: [1,294,0,15,5,0,0,255] redefinition_m2_subset_1
s: 294
i: [1,295,0,15,5,0,0,255] redefinition_m2_subset_1
s: 295
a: [3,1173,1,4,1,1160]
s: 1173
d: [2,1325,3,0,0,5,52,943,358,953,353,1173,291]
e: 1325
d: [2,1327,2,8,4,1,40,1173,291]
d: [2,1331,2,6,3,2,21,1327,353]
d: [2,1332,2,4,2,3,21,1331,953]
d: [2,1333,2,2,1,4,21,1332,358]
d: [2,1334,2,0,0,5,21,1333,943]
e: 1334
r: 1013
r: 1173
a: [3,1172,1,4,1,1160]
s: 1172
a: [3,1200,1,4,1,1187]
s: 1200
d: [2,754,1,2,1,1,40,257,445]
a: [3,765,1,2,1,754]
s: 765
d: [2,755,1,2,1,1,40,257,565]
a: [3,770,1,2,1,755]
s: 770
a: [3,1012,2,4,1,904]
s: 1012
d: [2,1365,3,0,0,5,52,902,696,853,701,1012,297]
e: 1365
d: [2,1366,3,8,4,1,40,1012,297]
d: [2,1371,3,6,3,2,21,1366,701]
d: [2,1372,3,4,2,3,21,1371,853]
d: [2,1373,3,2,1,4,21,1372,696]
d: [2,1374,3,0,0,5,21,1373,902]
e: 1374
d: [2,820,1,2,1,1,40,260,587]
d: [2,1048,2,4,2,3,40,410,843]
d: [2,1049,2,2,1,4,21,1048,717]
d: [2,1118,2,2,1,1,40,760,241]
d: [2,1119,2,2,1,1,40,760,240]
d: [2,1191,2,8,4,1,40,251,1149]
d: [2,1211,2,6,3,2,21,1191,353]
d: [2,1217,1,8,4,1,40,252,348]
d: [2,1226,1,6,3,2,21,1217,358]
d: [2,1232,2,4,2,1,40,1196,273]
d: [2,1234,2,2,1,2,21,1232,581]
d: [2,1244,3,2,1,2,21,1238,717]
d: [2,1254,1,16,5,1,40,248,338]
d: [2,1262,1,14,4,2,21,1254,1195]
d: [2,1338,1,10,3,3,21,1262,1168]
f: 1374,1338,1334,1292,1244,1234,1226,1211,1145,1119,1118,1049,1031,849,820,709,708,688,687,583,577,334,300,232,231,229,223,224,225,226,227
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t140_xreal_1
% SZS output start Proof for xreal_1__t140_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(k6_numbers,X0)))) [input t140_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(k6_numbers,X0)))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
12. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
15. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
22. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v3_xxreal_0(k7_xcmplx_0(X0,X1))) [input fc34_xreal_0]
30. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
38. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0))))) [input t6_real]
40. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
51. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
54. ! [X0,X1] : ((v1_xcmplx_0(X1) & ~v1_xboole_0(X1) & v1_xcmplx_0(X0) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k7_xcmplx_0(X0,X1))) [input fc9_xcmplx_0]
59. ? [X0] : (v1_xxreal_0(X0) & v1_xboole_0(X0)) [input rc4_xxreal_0]
60. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
73. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
79. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
80. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => r1_xxreal_0(X0,X0)) [input reflexivity_r1_xxreal_0]
81. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
82. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
85. ! [X0] : (v1_xreal_0(X0) => v1_xcmplx_0(X0)) [input cc3_xreal_0]
86. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
87. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k7_xcmplx_0(X0,X1))) [input fc8_xreal_0]
100. ? [X0] : (? [X1] : ((r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(k6_numbers,X0)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
101. ? [X0] : (? [X1] : (r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 100]
105. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
106. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
110. ! [X0] : (v1_xcmplx_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 85]
111. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 86]
112. ! [X0] : (! [X1] : (((v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0)) | ~r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 38]
113. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 112]
118. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 30]
119. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 118]
120. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 40]
121. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 120]
131. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 79]
132. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 51]
143. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 15]
144. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 143]
145. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 12]
146. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 145]
163. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 22]
164. ! [X0,X1] : (~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 163]
173. ! [X0,X1] : (~v1_xboole_0(k7_xcmplx_0(X0,X1)) | (~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0))) [ennf transformation 54]
174. ! [X0,X1] : (~v1_xboole_0(k7_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0)) [flattening 173]
177. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 87]
178. ! [X0,X1] : (v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 177]
179. ! [X0,X1] : (r1_xxreal_0(X0,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 80]
180. ! [X0,X1] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 179]
181. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 81]
182. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 181]
183. ? [X0] : (? [X1] : (r1_xxreal_0(k7_xcmplx_0(X0,X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(k6_numbers,X0) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (r1_xxreal_0(k7_xcmplx_0(sK0,X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
184. ? [X1] : (r1_xxreal_0(k7_xcmplx_0(sK0,X1),k6_numbers) & ~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(X1)) => (r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) & ~r1_xxreal_0(k6_numbers,sK1) & ~r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) [choice axiom]
185. (r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) & ~r1_xxreal_0(k6_numbers,sK1) & ~r1_xxreal_0(k6_numbers,sK0) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 101,184,183]
221. ? [X0] : (v1_xxreal_0(X0) & v1_xboole_0(X0)) => (v1_xxreal_0(sK19) & v1_xboole_0(sK19)) [choice axiom]
222. v1_xxreal_0(sK19) & v1_xboole_0(sK19) [skolemisation 59,221]
223. v1_xreal_0(sK0) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
224. v1_xreal_0(sK1) [cnf transformation 185] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
225. ~r1_xxreal_0(k6_numbers,sK0) [cnf transformation 185] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
226. ~r1_xxreal_0(k6_numbers,sK1) [cnf transformation 185] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
227. r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) [cnf transformation 185] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
229. v1_xboole_0(k1_xboole_0) [cnf transformation 73] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
231. k6_numbers = k1_xboole_0 [cnf transformation 82] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
232. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 60] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
240. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 105] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
241. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 106] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
246. ~v1_xreal_0(X0) | v1_xcmplx_0(X0) [cnf transformation 110] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
247. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 111] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
248. ~r1_xxreal_0(X0,X1) | v2_xxreal_0(X1) | v1_xboole_0(X0) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 113] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
251. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 119] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
252. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 121] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
259. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 131] {a:0,w:5,wCS:2000,sine:2,thAx:0,allAx:1,thDist:-1}
260. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 132] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
272. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 144] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
273. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 146] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
291. ~v3_xxreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 164] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
297. ~v1_xboole_0(k7_xcmplx_0(X0,X1)) | ~v1_xcmplx_0(X1) | v1_xboole_0(X1) | ~v1_xcmplx_0(X0) | v1_xboole_0(X0) [cnf transformation 174] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
299. v1_xreal_0(k7_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 178] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
300. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 180] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
301. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 182] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
334. v1_xxreal_0(sK19) [cnf transformation 222] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
336. 1 <=> r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) [avatar definition]
338. r1_xxreal_0(k7_xcmplx_0(sK0,sK1),k6_numbers) <- (1) [avatar component clause 336] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
339. 1 [avatar split clause 227,336]
341. 2 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
343. ~r1_xxreal_0(k6_numbers,sK1) <- (~2) [avatar component clause 341] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
344. ~2 [avatar split clause 226,341]
346. 3 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
348. ~r1_xxreal_0(k6_numbers,sK0) <- (~3) [avatar component clause 346] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
349. ~3 [avatar split clause 225,346]
351. 4 <=> v1_xreal_0(sK1) [avatar definition]
353. v1_xreal_0(sK1) <- (4) [avatar component clause 351] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
354. 4 [avatar split clause 224,351]
356. 5 <=> v1_xreal_0(sK0) [avatar definition]
358. v1_xreal_0(sK0) <- (5) [avatar component clause 356] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
359. 5 [avatar split clause 223,356]
366. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
368. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 366] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
369. 7 [avatar split clause 229,366]
376. 9 <=> k6_numbers = k1_xboole_0 [avatar definition]
378. k6_numbers = k1_xboole_0 <- (9) [avatar component clause 376] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
379. 9 [avatar split clause 231,376]
381. 10 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
383. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (10) [avatar component clause 381] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
384. 10 [avatar split clause 232,381]
406. 15 <=> ! [X1] : ~v1_xxreal_0(X1) [avatar definition]
407. ~v1_xxreal_0(X1) <- (15) [avatar component clause 406] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
409. 16 <=> ! [X0] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0)) [avatar definition]
410. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) <- (16) [avatar component clause 409] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
411. 15 | 16 [avatar split clause 300,409,406]
568. 48 <=> v1_xxreal_0(sK19) [avatar definition]
570. v1_xxreal_0(sK19) <- (48) [avatar component clause 568] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
571. 48 [avatar split clause 334,568]
577. v1_xboole_0(k6_numbers) <- (7, 9) [backward demodulation 368,378] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
579. 50 <=> v1_xboole_0(k6_numbers) [avatar definition]
581. v1_xboole_0(k6_numbers) <- (50) [avatar component clause 579] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
582. 50 | ~7 | ~9 [avatar split clause 577,376,366,579]
583. m1_subset_1(k6_numbers,k4_ordinal1) <- (9, 10) [forward demodulation 383,378] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
585. 51 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
587. m1_subset_1(k6_numbers,k4_ordinal1) <- (51) [avatar component clause 585] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
588. 51 | ~9 | ~10 [avatar split clause 583,381,376,585]
687. v1_xcmplx_0(sK0) <- (5) [resolution 246,358] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
688. v1_xcmplx_0(sK1) <- (4) [resolution 246,353] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
694. 68 <=> v1_xcmplx_0(sK0) [avatar definition]
696. v1_xcmplx_0(sK0) <- (68) [avatar component clause 694] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
697. 68 | ~5 [avatar split clause 687,356,694]
699. 69 <=> v1_xcmplx_0(sK1) [avatar definition]
701. v1_xcmplx_0(sK1) <- (69) [avatar component clause 699] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
702. 69 | ~4 [avatar split clause 688,351,699]
708. v1_xxreal_0(sK0) <- (5) [resolution 247,358] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
709. v1_xxreal_0(sK1) <- (4) [resolution 247,353] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
715. 71 <=> v1_xxreal_0(sK0) [avatar definition]
717. v1_xxreal_0(sK0) <- (71) [avatar component clause 715] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
718. 71 | ~5 [avatar split clause 708,356,715]
720. 72 <=> v1_xxreal_0(sK1) [avatar definition]
722. v1_xxreal_0(sK1) <- (72) [avatar component clause 720] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
723. 72 | ~4 [avatar split clause 709,351,720]
758. 78 <=> v7_ordinal1(k6_numbers) [avatar definition]
760. v7_ordinal1(k6_numbers) <- (78) [avatar component clause 758] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
819. k6_numbers = X0 | ~v1_xboole_0(X0) <- (9) [forward demodulation 259,378] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
820. v7_ordinal1(k6_numbers) <- (51) [resolution 260,587] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
822. 78 | ~51 [avatar split clause 820,585,758]
842. ~r1_xxreal_0(X0,sK1) | ~v1_xboole_0(X0) <- (~2, 9) [superposition 343,819] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
843. ~r1_xxreal_0(X0,sK0) | ~v1_xboole_0(X0) <- (~3, 9) [superposition 348,819] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
848. ~v1_xboole_0(sK1) | ~v1_xxreal_0(sK1) <- (~2, 9, 16) [resolution 842,410] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
849. ~v1_xboole_0(sK1) <- (~2, 9, 16, 72) [subsumption resolution 848,722] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
851. 91 <=> v1_xboole_0(sK1) [avatar definition]
853. ~v1_xboole_0(sK1) <- (~91) [avatar component clause 851] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
854. ~91 | 2 | ~9 | ~16 | ~72 [avatar split clause 849,720,409,376,341,851]
900. 98 <=> v1_xboole_0(sK0) [avatar definition]
902. ~v1_xboole_0(sK0) <- (~98) [avatar component clause 900] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
942. 101 <=> v2_xxreal_0(sK0) [avatar definition]
943. ~v2_xxreal_0(sK0) <- (~101) [avatar component clause 942] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
946. 102 <=> v3_xxreal_0(sK0) [avatar definition]
948. v3_xxreal_0(sK0) <- (102) [avatar component clause 946] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
952. 103 <=> v2_xxreal_0(sK1) [avatar definition]
953. ~v2_xxreal_0(sK1) <- (~103) [avatar component clause 952] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1011. 116 <=> v1_xboole_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1012. v1_xboole_0(k7_xcmplx_0(sK0,sK1)) <- (116) [avatar component clause 1011] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1031. $false <- (15, 48) [resolution 407,570] {a:1,w:0,sine:255,thAx:0,allAx:3,thDist:-3}
1032. ~15 | ~48 [avatar contradiction clause 1031]
1048. ~v1_xxreal_0(sK0) | ~v1_xboole_0(sK0) <- (~3, 9, 16) [resolution 410,843] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1049. ~v1_xboole_0(sK0) <- (~3, 9, 16, 71) [subsumption resolution 1048,717] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1050. ~98 | 3 | ~9 | ~16 | ~71 [avatar split clause 1049,715,409,376,346,900]
1118. v1_xxreal_0(k6_numbers) <- (78) [resolution 760,241] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1119. v1_xreal_0(k6_numbers) <- (78) [resolution 760,240] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1124. 120 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1127. 120 | ~78 [avatar split clause 1118,758,1124]
1129. 121 <=> v1_xreal_0(k6_numbers) [avatar definition]
1132. 121 | ~78 [avatar split clause 1119,758,1129]
1139. r1_xxreal_0(sK1,k6_numbers) | ~v1_xxreal_0(k6_numbers) | ~v1_xxreal_0(sK1) <- (~2) [resolution 301,343] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1145. r1_xxreal_0(sK1,k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (~2, 72) [subsumption resolution 1139,722] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1147. 123 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
1149. r1_xxreal_0(sK1,k6_numbers) <- (123) [avatar component clause 1147] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1150. ~120 | 123 | 2 | ~72 [avatar split clause 1145,720,341,1147,1124]
1167. 125 <=> v1_xreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1168. v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (125) [avatar component clause 1167] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1169. ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (~125) [avatar component clause 1167] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1171. 126 <=> v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) [avatar definition]
1173. v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) <- (126) [avatar component clause 1171] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1191. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (123) [resolution 251,1149] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1194. 129 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1195. ~v2_xxreal_0(k6_numbers) <- (~129) [avatar component clause 1194] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1196. v2_xxreal_0(k6_numbers) <- (129) [avatar component clause 1194] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1211. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (4, 123) [subsumption resolution 1191,353] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1212. ~121 | 129 | ~103 | ~4 | ~123 [avatar split clause 1211,1147,351,952,1194,1129]
1217. v3_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (~3) [resolution 252,348] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1226. v3_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (~3, 5) [subsumption resolution 1217,358] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1227. ~121 | 129 | 102 | 3 | ~5 [avatar split clause 1226,356,346,946,1194,1129]
1232. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (129) [resolution 1196,273] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1234. ~v1_xxreal_0(k6_numbers) <- (50, 129) [subsumption resolution 1232,581] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1235. ~120 | ~50 | ~129 [avatar split clause 1234,1194,579,1124]
1238. ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (102) [resolution 948,272] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1244. ~v2_xxreal_0(sK0) <- (71, 102) [subsumption resolution 1238,717] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1245. ~101 | ~71 | ~102 [avatar split clause 1244,946,715,942]
1254. v2_xxreal_0(k6_numbers) | v1_xboole_0(k7_xcmplx_0(sK0,sK1)) | v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (1) [resolution 248,338] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1262. v1_xboole_0(k7_xcmplx_0(sK0,sK1)) | v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(k7_xcmplx_0(sK0,sK1)) <- (1, ~129) [subsumption resolution 1254,1195] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1289. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~125) [resolution 1169,299] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1291. ~v1_xreal_0(sK0) <- (4, ~125) [subsumption resolution 1289,353] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1292. $false <- (4, 5, ~125) [subsumption resolution 1291,358] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1293. ~4 | ~5 | 125 [avatar contradiction clause 1292]
1327. ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (126) [resolution 1173,291] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1331. v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (4, 126) [subsumption resolution 1327,353] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1332. ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (4, ~103, 126) [subsumption resolution 1331,953] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1333. v2_xxreal_0(sK0) <- (4, 5, ~103, 126) [subsumption resolution 1332,358] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1334. $false <- (4, 5, ~101, ~103, 126) [subsumption resolution 1333,943] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1335. ~4 | ~5 | 101 | 103 | ~126 [avatar contradiction clause 1334]
1338. v1_xboole_0(k7_xcmplx_0(sK0,sK1)) | v3_xxreal_0(k7_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (1, 125, ~129) [subsumption resolution 1262,1168] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1339. ~121 | 126 | 116 | ~1 | ~125 | 129 [avatar split clause 1338,1194,1167,336,1011,1171,1129]
1366. ~v1_xcmplx_0(sK1) | v1_xboole_0(sK1) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (116) [resolution 1012,297] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1371. v1_xboole_0(sK1) | ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (69, 116) [subsumption resolution 1366,701] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1372. ~v1_xcmplx_0(sK0) | v1_xboole_0(sK0) <- (69, ~91, 116) [subsumption resolution 1371,853] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1373. v1_xboole_0(sK0) <- (68, 69, ~91, 116) [subsumption resolution 1372,696] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1374. $false <- (68, 69, ~91, ~98, 116) [subsumption resolution 1373,902] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
1375. ~68 | ~69 | 91 | 98 | ~116 [avatar contradiction clause 1374]
1376. $false [avatar sat refutation 339,344,349,354,359,369,379,384,411,571,582,588,697,702,718,723,822,854,1032,1050,1127,1132,1150,1212,1227,1235,1245,1293,1335,1339,1375] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1028,thDist:-1028}
% SZS output end Proof for xreal_1__t140_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 87
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 112
% Generated clauses: 575
% Active clauses: 155
% Passive clauses: 237
% Final active clauses: 140
% Final passive clauses: 67
% Main loop iterations started: 167
% 
% >>> Simplifying Inferences
% Duplicate literals: 21
% Fw subsumption resolutions: 62
% Fw demodulations: 4
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Forward subsumptions: 68
% Backward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 199
% Unit resulting resolution: 7
% Factoring: 1
% Forward superposition: 15
% Backward superposition: 6
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 37
% Split components: 96
% Sat splitting refutations: 23
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 184
% SAT solver unit clauses: 50
% SAT solver binary clauses: 66
% 
% Memory used [KB]: 5756
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% splitting model update: 0.001 s ( own 0.0000 s ) 
% unit clause index maintenance: 0.001 s
% backward superposition index maintenance: 0.001 s
% other: 0.006 s ( own 0.003 s ) 

