i: [1,42,0,3,1,1,0,0] t3_zfmisc_1
a: [3,69,0,3,1,42]
s: 69
i: [1,41,0,5,1,1,0,0] t3_zfmisc_1
a: [3,74,0,5,1,41]
s: 74
i: [1,43,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,79,0,2,1,43]
s: 79
i: [1,44,0,2,1,0,0,2] fc1_xboole_0
a: [3,84,0,2,1,44]
s: 84
i: [1,59,0,2,1,0,0,2] rc2_xboole_0
a: [3,94,0,2,1,59]
s: 94
i: [1,60,0,2,1,0,0,2] rc1_xboole_0
a: [3,99,0,2,1,60]
s: 99
i: [1,46,0,3,1,0,0,1] fc2_xboole_0
s: 46
i: [1,57,0,3,1,0,0,1] reflexivity_r1_tarski
s: 57
i: [1,45,0,3,1,0,0,255] d2_xboole_0
a: [3,89,0,3,1,45]
s: 89
r: 79
i: [1,62,0,4,1,0,0,255] d1_tarski
s: 62
i: [1,65,0,4,1,0,0,1] l3_zfmisc_1
s: 65
i: [1,47,0,5,2,0,0,2] t6_boole
s: 47
i: [1,56,0,5,2,0,0,255] t7_boole
s: 56
i: [1,64,0,5,1,0,0,1] l3_zfmisc_1
s: 64
i: [1,58,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 58
i: [1,55,0,7,3,0,0,1] t8_boole
s: 55
d: [2,106,1,4,1,0,40,58,62]
s: 106
i: [1,63,0,7,2,0,0,255] d1_tarski
s: 63
d: [2,103,1,6,2,0,44,65,47]
s: 103
i: [1,52,0,11,3,0,0,1] l3_zfmisc_1
s: 52
i: [1,50,0,14,3,0,0,255] d1_tarski
s: 50
d: [2,108,1,9,2,1,40,52,74]
a: [3,119,1,4,1,108]
s: 119
r: 74
d: [2,125,2,2,1,1,44,46,119]
d: [2,133,2,0,0,2,21,125,84]
e: 133
r: 119
s: 74
a: [3,118,1,4,1,108]
s: 118
i: [1,51,0,14,3,0,0,255] d1_tarski
s: 51
a: [3,115,1,5,1,108]
s: 115
r: 74
d: [2,140,2,4,1,1,44,62,115]
a: [3,148,2,4,1,140]
s: 148
d: [2,155,3,0,0,2,52,69,148,63]
e: 155
d: [2,156,3,3,1,1,40,148,63]
d: [2,160,3,0,0,2,21,156,69]
e: 160
d: [2,135,1,9,2,1,40,74,52]
f: 160,140,135,133,44,41,42
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t3_zfmisc_1
% SZS output start Proof for zfmisc_1__t3_zfmisc_1
1. ! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),k1_tarski(X1)) => X0 = X1) [input t3_zfmisc_1]
2. ~! [X0] : ! [X1] : (r1_tarski(k1_tarski(X0),k1_tarski(X1)) => X0 = X1) [negated conjecture 1]
4. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
10. ! [X0] : ~v1_xboole_0(k1_tarski(X0)) [input fc2_xboole_0]
11. ! [X0] : ! [X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [input l3_zfmisc_1]
18. ~! [X0,X1] : (r1_tarski(k1_tarski(X0),k1_tarski(X1)) => X0 = X1) [flattening 2]
19. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 4]
20. ! [X0,X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [flattening 11]
24. ? [X0,X1] : (X0 != X1 & r1_tarski(k1_tarski(X0),k1_tarski(X1))) [ennf transformation 18]
29. ? [X0,X1] : (X0 != X1 & r1_tarski(k1_tarski(X0),k1_tarski(X1))) => (sK0 != sK1 & r1_tarski(k1_tarski(sK0),k1_tarski(sK1))) [choice axiom]
30. sK0 != sK1 & r1_tarski(k1_tarski(sK0),k1_tarski(sK1)) [skolemisation 24,29]
31. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 19]
32. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 31]
33. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) [choice axiom]
34. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 32,33]
35. ! [X0,X1] : ((r1_tarski(X0,k1_tarski(X1)) | (k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & ((k1_tarski(X1) = X0 | k1_xboole_0 = X0) | ~r1_tarski(X0,k1_tarski(X1)))) [nnf transformation 20]
36. ! [X0,X1] : ((r1_tarski(X0,k1_tarski(X1)) | (k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & (k1_tarski(X1) = X0 | k1_xboole_0 = X0 | ~r1_tarski(X0,k1_tarski(X1)))) [flattening 35]
41. r1_tarski(k1_tarski(sK0),k1_tarski(sK1)) [cnf transformation 30] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. sK0 != sK1 [cnf transformation 30] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
44. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
46. ~v1_xboole_0(k1_tarski(X0)) [cnf transformation 10] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
48. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 34] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
49. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 34] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
52. ~r1_tarski(X0,k1_tarski(X1)) | k1_xboole_0 = X0 | k1_tarski(X1) = X0 [cnf transformation 36] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 49] {a:0,w:7,wCS:2800,sine:255,thAx:0,allAx:2,thDist:-2}
62. r2_hidden(X3,k1_tarski(X3)) [equality resolution 61] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 48] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
67. 1 <=> sK0 = sK1 [avatar definition]
69. sK0 != sK1 <- (~1) [avatar component clause 67] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
70. ~1 [avatar split clause 42,67]
72. 2 <=> r1_tarski(k1_tarski(sK0),k1_tarski(sK1)) [avatar definition]
74. r1_tarski(k1_tarski(sK0),k1_tarski(sK1)) <- (2) [avatar component clause 72] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
75. 2 [avatar split clause 41,72]
82. 4 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
84. v1_xboole_0(k1_xboole_0) <- (4) [avatar component clause 82] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
85. 4 [avatar split clause 44,82]
113. 8 <=> k1_tarski(sK0) = k1_tarski(sK1) [avatar definition]
115. k1_tarski(sK0) = k1_tarski(sK1) <- (8) [avatar component clause 113] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
117. 9 <=> k1_xboole_0 = k1_tarski(sK0) [avatar definition]
119. k1_xboole_0 = k1_tarski(sK0) <- (9) [avatar component clause 117] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
125. ~v1_xboole_0(k1_xboole_0) <- (9) [superposition 46,119] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
133. $false <- (4, 9) [subsumption resolution 125,84] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
134. ~4 | ~9 [avatar contradiction clause 133]
135. k1_xboole_0 = k1_tarski(sK0) | k1_tarski(sK0) = k1_tarski(sK1) <- (2) [resolution 74,52] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
136. 8 | 9 | ~2 [avatar split clause 135,72,117,113]
140. r2_hidden(sK1,k1_tarski(sK0)) <- (8) [superposition 62,115] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
146. 10 <=> r2_hidden(sK1,k1_tarski(sK0)) [avatar definition]
148. r2_hidden(sK1,k1_tarski(sK0)) <- (10) [avatar component clause 146] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
149. 10 | ~8 [avatar split clause 140,113,146]
156. sK0 = sK1 <- (10) [resolution 148,63] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
160. $false <- (~1, 10) [subsumption resolution 156,69] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
161. 1 | ~10 [avatar contradiction clause 160]
162. $false [avatar sat refutation 70,75,85,134,136,149,161] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:65,thDist:-65}
% SZS output end Proof for zfmisc_1__t3_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 20
% Generated clauses: 71
% Active clauses: 26
% Passive clauses: 33
% Final active clauses: 22
% Final passive clauses: 6
% Main loop iterations started: 27
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 2
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Equational tautologies: 3
% Forward subsumptions: 12
% 
% >>> Generating Inferences
% Binary resolution: 14
% Unit resulting resolution: 1
% Backward superposition: 16
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 7
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% backward demodulation: 0.001 s ( own 0.0000 s ) 
% other: 0.001 s ( own 0.0000 s ) 

