i: [1,145,0,5,2,1,0,0] t9_grfunc_1
a: [3,223,0,2,1,145]
s: 223
i: [1,141,0,2,1,1,0,0] t9_grfunc_1
a: [3,246,0,2,1,141]
s: 246
i: [1,140,0,2,1,1,0,0] t9_grfunc_1
a: [3,251,0,2,1,140]
s: 251
i: [1,146,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,256,0,2,1,146]
s: 256
i: [1,199,0,2,1,0,0,4] rc7_funct_1
a: [3,266,0,2,1,199]
s: 266
i: [1,198,0,2,1,0,0,4] rc7_funct_1
a: [3,271,0,2,1,198]
s: 271
i: [1,201,0,2,1,0,0,2] rc1_relat_1
a: [3,276,0,2,1,201]
s: 276
i: [1,200,0,2,1,0,0,2] rc1_relat_1
a: [3,281,0,2,1,200]
s: 281
i: [1,205,0,2,1,0,0,3] rc4_funct_1
a: [3,286,0,2,1,205]
s: 286
i: [1,204,0,2,1,0,0,3] rc4_funct_1
a: [3,291,0,2,1,204]
s: 291
i: [1,203,0,2,1,0,0,3] rc4_funct_1
a: [3,296,0,2,1,203]
s: 296
i: [1,202,0,2,1,0,0,3] rc4_funct_1
a: [3,301,0,2,1,202]
s: 301
i: [1,207,0,2,1,0,0,3] rc2_relat_1
a: [3,306,0,2,1,207]
s: 306
i: [1,206,0,2,1,0,0,3] rc2_relat_1
a: [3,311,0,2,1,206]
s: 311
i: [1,210,0,2,1,0,0,3] rc3_funct_1
a: [3,316,0,2,1,210]
s: 316
i: [1,209,0,2,1,0,0,3] rc3_funct_1
a: [3,321,0,2,1,209]
s: 321
i: [1,208,0,2,1,0,0,3] rc3_funct_1
a: [3,326,0,2,1,208]
s: 326
i: [1,212,0,2,1,0,0,1] rc1_funct_1
a: [3,331,0,2,1,212]
s: 331
i: [1,211,0,2,1,0,0,1] rc1_funct_1
a: [3,336,0,2,1,211]
s: 336
i: [1,215,0,2,1,0,0,255] rc5_funct_1
a: [3,341,0,2,1,215]
s: 341
i: [1,214,0,2,1,0,0,255] rc5_funct_1
a: [3,346,0,2,1,214]
s: 346
i: [1,213,0,2,1,0,0,255] rc5_funct_1
a: [3,351,0,2,1,213]
s: 351
i: [1,218,0,2,1,0,0,1] rc2_funct_1
a: [3,356,0,2,1,218]
s: 356
i: [1,217,0,2,1,0,0,1] rc2_funct_1
a: [3,361,0,2,1,217]
s: 361
i: [1,216,0,2,1,0,0,1] rc2_funct_1
a: [3,366,0,2,1,216]
s: 366
i: [1,147,0,3,1,0,0,255] d2_xboole_0
a: [3,261,0,3,1,147]
s: 261
r: 256
d: [2,368,0,2,1,2,23,256,261]
a: [3,372,0,2,1,368]
s: 372
i: [1,142,0,15,4,1,0,0] t9_grfunc_1
a: [3,241,0,13,3,142]
s: 241
i: [1,150,0,4,2,0,0,4] cc7_funct_1
s: 150
i: [1,151,0,4,2,0,0,1] cc1_funct_1
s: 151
i: [1,152,0,4,2,0,0,2] cc1_relat_1
s: 152
d: [2,374,1,2,1,1,40,150,372]
a: [3,378,1,2,1,374]
s: 378
i: [1,189,0,4,1,0,0,3] existence_m1_subset_1
s: 189
d: [2,380,1,2,1,1,40,151,372]
a: [3,384,1,2,1,380]
s: 384
i: [1,153,0,5,2,0,0,255] t6_boole
s: 153
i: [1,154,0,5,2,0,0,2] fc10_relat_1
s: 154
d: [2,386,1,2,1,1,40,152,372]
a: [3,390,1,2,1,386]
s: 390
i: [1,155,0,5,2,0,0,2] fc4_xtuple_0
s: 155
i: [1,185,0,5,2,0,0,1] t7_boole
s: 185
i: [1,190,0,5,1,0,0,255] fc1_funct_1
s: 190
i: [1,191,0,5,1,0,0,255] fc5_relat_1
s: 191
d: [2,396,1,5,2,0,40,154,152]
s: 396
i: [1,169,0,6,3,0,0,255] cc3_relat_1
s: 169
d: [2,400,0,4,2,0,21,169,152]
s: 400
i: [1,171,0,6,3,0,0,3] cc4_relat_1
s: 171
d: [2,408,0,4,2,0,21,171,152]
s: 408
d: [2,401,1,2,1,1,40,400,372]
a: [3,406,1,2,1,401]
s: 406
i: [1,182,0,6,2,0,0,3] t1_subset
s: 182
d: [2,409,1,2,1,1,40,408,372]
a: [3,414,1,2,1,409]
s: 414
i: [1,193,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 193
i: [1,148,0,7,3,0,0,3] cc8_funct_1
s: 148
d: [2,397,1,5,2,0,40,154,151]
s: 397
i: [1,149,0,7,3,0,0,3] cc8_funct_1
s: 149
d: [2,398,1,5,2,0,40,154,150]
s: 398
i: [1,157,0,7,3,0,0,2] fc8_relat_1
s: 157
i: [1,161,0,7,3,0,0,255] fc24_relat_1
s: 161
d: [2,402,1,5,2,0,40,400,154]
s: 402
i: [1,175,0,7,3,0,0,255] fc14_funct_1
s: 175
d: [2,410,1,5,2,0,40,408,154]
s: 410
i: [1,184,0,7,3,0,0,1] t8_boole
s: 184
i: [1,192,0,7,1,0,0,255] commutativity_k2_tarski
s: 192
d: [2,418,1,5,2,0,40,148,189]
s: 418
i: [1,160,0,8,4,0,0,255] cc6_funct_1
s: 160
d: [2,419,1,5,2,0,40,149,189]
s: 419
i: [1,164,0,8,4,0,0,1] cc2_funct_1
s: 164
d: [2,422,0,6,3,0,21,164,152]
d: [2,423,0,4,2,0,21,422,151]
s: 423
d: [2,424,1,2,1,1,40,423,372]
a: [3,429,1,2,1,424]
s: 429
i: [1,167,0,8,4,0,0,255] cc4_funct_1
s: 167
d: [2,432,0,6,3,0,21,167,152]
d: [2,433,0,4,2,0,21,432,151]
s: 433
i: [1,172,0,8,4,0,0,255] cc5_funct_1
s: 172
i: [1,183,0,8,3,0,0,3] t2_subset
s: 183
d: [2,434,1,2,1,1,40,433,372]
a: [3,439,1,2,1,434]
s: 439
i: [1,197,0,8,1,0,0,255] fc7_relat_1
s: 197
d: [2,425,1,5,2,0,40,423,154]
s: 425
i: [1,156,0,9,4,0,0,255] fc18_funct_1
s: 156
i: [1,181,0,10,1,0,0,255] d5_tarski
s: 181
d: [2,435,1,5,2,0,40,433,154]
s: 435
i: [1,195,0,10,4,0,0,255] fc19_funct_1
s: 195
d: [2,441,1,6,2,0,40,183,189]
s: 441
i: [1,177,0,11,4,0,0,1] d4_funct_1
s: 177
i: [1,178,0,11,4,0,0,1] d4_funct_1
s: 178
d: [2,461,2,6,2,0,40,441,193]
s: 461
i: [1,180,0,11,4,0,0,1] d4_funct_1
s: 180
d: [2,392,1,7,3,0,44,153,153]
s: 392
i: [1,196,0,12,5,0,0,255] fc15_funct_1
s: 196
i: [1,186,0,13,4,0,0,2] t1_funct_1
s: 186
d: [2,466,1,9,4,0,40,177,185]
s: 466
i: [1,187,0,14,4,0,0,2] t1_funct_1
s: 187
d: [2,469,1,9,4,0,40,178,185]
s: 469
i: [1,179,0,15,4,0,0,1] d4_funct_1
s: 179
i: [1,219,0,15,4,0,0,2] t1_funct_1
s: 219
d: [2,476,1,18,5,1,40,219,241]
d: [2,486,1,16,4,2,21,476,246]
d: [2,487,1,14,3,3,21,486,251]
s: 487
i: [1,194,0,16,6,0,0,2] fc11_funct_1
s: 194
d: [2,492,2,23,6,3,44,487,179]
d: [2,494,2,21,5,4,21,492,223]
d: [2,495,2,19,4,4,21,494,246]
d: [2,496,2,17,3,4,21,495,251]
a: [3,500,2,5,1,496]
s: 500
d: [2,506,3,0,0,4,52,251,246,223,500,178]
e: 506
d: [2,507,3,6,3,1,40,500,178]
d: [2,509,3,4,2,2,21,507,223]
d: [2,510,3,2,1,3,21,509,246]
d: [2,511,3,0,0,4,21,510,251]
e: 511
r: 500
i: [1,176,0,24,7,0,0,1] d4_funct_1
s: 176
d: [2,481,1,10,4,0,40,219,185]
d: [2,488,1,8,3,0,21,481,152]
d: [2,489,1,6,2,0,21,488,151]
s: 489
a: [3,503,2,12,2,496]
s: 503
d: [2,447,1,10,1,0,44,181,192]
s: 447
d: [2,525,3,14,4,1,40,503,219]
d: [2,526,3,12,3,2,21,525,246]
d: [2,527,3,10,2,3,21,526,251]
a: [3,531,3,5,1,527]
s: 531
d: [2,547,4,0,0,4,52,251,246,223,531,177]
e: 547
d: [2,548,4,6,3,1,40,531,177]
d: [2,550,4,4,2,2,21,548,223]
d: [2,551,4,2,1,3,21,550,246]
d: [2,552,4,0,0,4,21,551,251]
e: 552
r: 531
d: [2,448,1,10,1,0,44,181,192]
s: 448
d: [2,450,1,10,1,0,44,181,192]
s: 450
a: [3,535,3,5,1,527]
s: 535
r: 503
d: [2,563,4,0,0,4,52,251,246,223,535,180]
e: 563
d: [2,564,4,11,4,1,44,180,535]
d: [2,566,4,6,3,1,19,564]
d: [2,568,4,4,2,2,21,566,223]
d: [2,569,4,2,1,3,21,568,246]
d: [2,570,4,0,0,4,21,569,251]
e: 570
r: 223
r: 535
a: [3,222,0,2,1,145]
s: 222
a: [3,227,0,3,1,145]
s: 227
i: [1,144,0,7,2,1,0,0] t9_grfunc_1
a: [3,232,0,5,1,144]
s: 232
i: [1,143,0,7,2,1,0,0] t9_grfunc_1
a: [3,237,0,5,1,143]
s: 237
d: [2,610,1,0,0,4,52,227,232,237,241]
e: 610
d: [2,611,1,0,0,4,52,227,232,237,241]
e: 611
r: 487
r: 241
d: [2,580,1,2,1,1,40,232,185]
a: [3,608,1,2,1,580]
s: 608
d: [2,577,1,8,3,1,40,232,186]
d: [2,588,1,6,2,2,21,577,246]
d: [2,589,1,4,1,3,21,588,251]
a: [3,593,1,4,1,589]
s: 593
d: [2,648,2,3,1,1,40,593,185]
a: [3,662,2,3,1,648]
s: 662
d: [2,614,1,8,3,1,40,237,186]
d: [2,622,1,6,2,2,21,614,246]
d: [2,623,1,4,1,3,21,622,251]
a: [3,638,1,4,1,623]
s: 638
d: [2,576,1,9,3,1,40,232,187]
d: [2,581,1,7,2,2,21,576,246]
d: [2,582,1,5,1,3,21,581,251]
a: [3,586,1,5,1,582]
s: 586
d: [2,682,2,8,4,1,44,195,586]
d: [2,691,2,6,3,2,21,682,246]
d: [2,692,2,4,2,3,21,691,251]
a: [3,696,2,2,1,692]
s: 696
d: [2,578,1,5,1,1,40,232,193]
a: [3,598,1,5,1,578]
s: 598
d: [2,646,2,4,1,1,40,593,193]
a: [3,652,2,4,1,646]
s: 652
d: [2,579,1,5,1,1,40,232,182]
a: [3,603,1,5,1,579]
s: 603
d: [2,616,1,5,1,1,40,237,182]
a: [3,628,1,5,1,616]
s: 628
d: [2,703,2,6,2,1,40,603,149]
a: [3,708,2,2,1,703]
s: 708
d: [2,615,1,5,1,1,40,237,193]
a: [3,633,1,5,1,615]
s: 633
d: [2,647,2,4,1,1,40,593,182]
a: [3,657,2,4,1,647]
s: 657
d: [2,613,1,9,3,1,40,237,187]
d: [2,620,1,7,2,2,21,613,246]
d: [2,621,1,5,1,3,21,620,251]
a: [3,643,1,5,1,621]
s: 643
d: [2,451,1,10,1,0,44,181,192]
s: 451
a: [3,699,2,2,1,692]
s: 699
d: [2,452,1,10,1,0,44,192,181]
s: 452
d: [2,734,3,5,2,1,40,657,149]
a: [3,748,3,3,1,734]
s: 748
d: [2,453,1,10,1,0,44,192,181]
s: 453
d: [2,454,1,10,4,0,40,195,150]
s: 454
d: [2,666,2,4,1,1,40,638,193]
a: [3,672,2,4,1,666]
s: 672
d: [2,455,1,10,4,0,40,195,151]
s: 455
d: [2,667,2,4,1,1,40,638,182]
a: [3,677,2,4,1,667]
s: 677
d: [2,456,1,10,4,0,40,195,152]
s: 456
d: [2,457,1,10,4,0,40,195,400]
s: 457
d: [2,679,2,22,7,1,44,176,586]
d: [2,683,2,18,6,2,21,679,593]
d: [2,684,2,16,5,3,21,683,222]
d: [2,685,2,14,4,4,21,684,246]
d: [2,686,2,12,3,5,21,685,251]
s: 686
d: [2,796,3,0,0,8,52,227,638,643,686]
e: 796
d: [2,798,3,10,3,6,44,686,643]
d: [2,800,3,7,2,6,19,798]
d: [2,802,3,4,1,7,21,800,227]
d: [2,803,3,0,0,8,21,802,638]
e: 803
d: [2,572,2,21,5,3,21,492,246]
d: [2,573,2,19,4,3,21,572,251]
f: 803,621,623,589,582,573,570,552,527,511,140,141,142,143,144,145
% Refutation found. Thanks to Tanya!
% SZS status Theorem for grfunc_1__t9_grfunc_1
% SZS output start Proof for grfunc_1__t9_grfunc_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) <=> ! [X1] : ! [X2] : ! [X3] : ((r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(k4_tarski(X1,X3),X0)) => X1 = X2))) [input t9_grfunc_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) <=> ! [X1] : ! [X2] : ! [X3] : ((r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(k4_tarski(X1,X3),X0)) => X1 = X2))) [negated conjecture 1]
16. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) <=> ! [X1] : ! [X2] : ((k1_funct_1(X0,X1) = k1_funct_1(X0,X2) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(X1,k9_xtuple_0(X0))) => X1 = X2))) [input d4_funct_1]
49. ! [X0] : ! [X1] : ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r2_hidden(k4_tarski(X0,X1),X2) <=> (k1_funct_1(X2,X0) = X1 & r2_hidden(X0,k9_xtuple_0(X2))))) [input t1_funct_1]
55. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) <=> ! [X1,X2,X3] : ((r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(k4_tarski(X1,X3),X0)) => X1 = X2))) [flattening 2]
56. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) <=> ! [X1,X2] : ((k1_funct_1(X0,X1) = k1_funct_1(X0,X2) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(X1,k9_xtuple_0(X0))) => X1 = X2))) [flattening 16]
62. ! [X0,X1,X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r2_hidden(k4_tarski(X0,X1),X2) <=> (k1_funct_1(X2,X0) = X1 & r2_hidden(X0,k9_xtuple_0(X2))))) [flattening 49]
65. ? [X0] : ((v2_funct_1(X0) <~> ! [X1,X2,X3] : (X1 = X2 | (~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(k4_tarski(X1,X3),X0)))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 55]
66. ? [X0] : ((v2_funct_1(X0) <~> ! [X1,X2,X3] : (X1 = X2 | ~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(k4_tarski(X1,X3),X0))) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 65]
94. ! [X0] : ((v2_funct_1(X0) <=> ! [X1,X2] : (X1 = X2 | (k1_funct_1(X0,X1) != k1_funct_1(X0,X2) | ~r2_hidden(X2,k9_xtuple_0(X0)) | ~r2_hidden(X1,k9_xtuple_0(X0))))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 56]
95. ! [X0] : ((v2_funct_1(X0) <=> ! [X1,X2] : (X1 = X2 | k1_funct_1(X0,X1) != k1_funct_1(X0,X2) | ~r2_hidden(X2,k9_xtuple_0(X0)) | ~r2_hidden(X1,k9_xtuple_0(X0)))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 94]
101. ! [X0,X1,X2] : ((r2_hidden(k4_tarski(X0,X1),X2) <=> (k1_funct_1(X2,X0) = X1 & r2_hidden(X0,k9_xtuple_0(X2)))) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) [ennf transformation 62]
102. ! [X0,X1,X2] : ((r2_hidden(k4_tarski(X0,X1),X2) <=> (k1_funct_1(X2,X0) = X1 & r2_hidden(X0,k9_xtuple_0(X2)))) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) [flattening 101]
110. ? [X0] : (((? [X1,X2,X3] : (X1 != X2 & r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(k4_tarski(X1,X3),X0)) | ~v2_funct_1(X0)) & (! [X1,X2,X3] : (X1 = X2 | ~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(k4_tarski(X1,X3),X0)) | v2_funct_1(X0))) & v1_funct_1(X0) & v1_relat_1(X0)) [nnf transformation 66]
111. ? [X0] : ((? [X1,X2,X3] : (X1 != X2 & r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(k4_tarski(X1,X3),X0)) | ~v2_funct_1(X0)) & (! [X1,X2,X3] : (X1 = X2 | ~r2_hidden(k4_tarski(X2,X3),X0) | ~r2_hidden(k4_tarski(X1,X3),X0)) | v2_funct_1(X0)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 110]
112. ? [X0] : ((? [X1,X2,X3] : (X1 != X2 & r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(k4_tarski(X1,X3),X0)) | ~v2_funct_1(X0)) & (! [X4,X5,X6] : (X4 = X5 | ~r2_hidden(k4_tarski(X5,X6),X0) | ~r2_hidden(k4_tarski(X4,X6),X0)) | v2_funct_1(X0)) & v1_funct_1(X0) & v1_relat_1(X0)) [rectify 111]
113. ? [X0] : ((? [X1,X2,X3] : (X1 != X2 & r2_hidden(k4_tarski(X2,X3),X0) & r2_hidden(k4_tarski(X1,X3),X0)) | ~v2_funct_1(X0)) & (! [X4,X5,X6] : (X4 = X5 | ~r2_hidden(k4_tarski(X5,X6),X0) | ~r2_hidden(k4_tarski(X4,X6),X0)) | v2_funct_1(X0)) & v1_funct_1(X0) & v1_relat_1(X0)) => ((? [X3,X2,X1] : (X1 != X2 & r2_hidden(k4_tarski(X2,X3),sK0) & r2_hidden(k4_tarski(X1,X3),sK0)) | ~v2_funct_1(sK0)) & (! [X6,X5,X4] : (X4 = X5 | ~r2_hidden(k4_tarski(X5,X6),sK0) | ~r2_hidden(k4_tarski(X4,X6),sK0)) | v2_funct_1(sK0)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
114. ? [X3,X2,X1] : (X1 != X2 & r2_hidden(k4_tarski(X2,X3),sK0) & r2_hidden(k4_tarski(X1,X3),sK0)) => (sK1 != sK2 & r2_hidden(k4_tarski(sK2,sK3),sK0) & r2_hidden(k4_tarski(sK1,sK3),sK0)) [choice axiom]
115. ((sK1 != sK2 & r2_hidden(k4_tarski(sK2,sK3),sK0) & r2_hidden(k4_tarski(sK1,sK3),sK0)) | ~v2_funct_1(sK0)) & (! [X4,X5,X6] : (X4 = X5 | ~r2_hidden(k4_tarski(X5,X6),sK0) | ~r2_hidden(k4_tarski(X4,X6),sK0)) | v2_funct_1(sK0)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 112,114,113]
116. ! [X0] : (((v2_funct_1(X0) | ? [X1,X2] : (X1 != X2 & k1_funct_1(X0,X1) = k1_funct_1(X0,X2) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(X1,k9_xtuple_0(X0)))) & (! [X1,X2] : (X1 = X2 | k1_funct_1(X0,X1) != k1_funct_1(X0,X2) | ~r2_hidden(X2,k9_xtuple_0(X0)) | ~r2_hidden(X1,k9_xtuple_0(X0))) | ~v2_funct_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [nnf transformation 95]
117. ! [X0] : (((v2_funct_1(X0) | ? [X1,X2] : (X1 != X2 & k1_funct_1(X0,X1) = k1_funct_1(X0,X2) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(X1,k9_xtuple_0(X0)))) & (! [X3,X4] : (X3 = X4 | k1_funct_1(X0,X3) != k1_funct_1(X0,X4) | ~r2_hidden(X4,k9_xtuple_0(X0)) | ~r2_hidden(X3,k9_xtuple_0(X0))) | ~v2_funct_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [rectify 116]
118. ! [X0] : (? [X1,X2] : (X1 != X2 & k1_funct_1(X0,X1) = k1_funct_1(X0,X2) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(X1,k9_xtuple_0(X0))) => (sK4(X0) != sK5(X0) & k1_funct_1(X0,sK4(X0)) = k1_funct_1(X0,sK5(X0)) & r2_hidden(sK5(X0),k9_xtuple_0(X0)) & r2_hidden(sK4(X0),k9_xtuple_0(X0)))) [choice axiom]
119. ! [X0] : (((v2_funct_1(X0) | (sK4(X0) != sK5(X0) & k1_funct_1(X0,sK4(X0)) = k1_funct_1(X0,sK5(X0)) & r2_hidden(sK5(X0),k9_xtuple_0(X0)) & r2_hidden(sK4(X0),k9_xtuple_0(X0)))) & (! [X3,X4] : (X3 = X4 | k1_funct_1(X0,X3) != k1_funct_1(X0,X4) | ~r2_hidden(X4,k9_xtuple_0(X0)) | ~r2_hidden(X3,k9_xtuple_0(X0))) | ~v2_funct_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [skolemisation 117,118]
120. ! [X0,X1,X2] : (((r2_hidden(k4_tarski(X0,X1),X2) | (k1_funct_1(X2,X0) != X1 | ~r2_hidden(X0,k9_xtuple_0(X2)))) & ((k1_funct_1(X2,X0) = X1 & r2_hidden(X0,k9_xtuple_0(X2))) | ~r2_hidden(k4_tarski(X0,X1),X2))) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) [nnf transformation 102]
121. ! [X0,X1,X2] : (((r2_hidden(k4_tarski(X0,X1),X2) | k1_funct_1(X2,X0) != X1 | ~r2_hidden(X0,k9_xtuple_0(X2))) & ((k1_funct_1(X2,X0) = X1 & r2_hidden(X0,k9_xtuple_0(X2))) | ~r2_hidden(k4_tarski(X0,X1),X2))) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) [flattening 120]
140. v1_relat_1(sK0) [cnf transformation 115] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
141. v1_funct_1(sK0) [cnf transformation 115] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
142. X4 = X5 | ~r2_hidden(k4_tarski(X5,X6),sK0) | ~r2_hidden(k4_tarski(X4,X6),sK0) | v2_funct_1(sK0) [cnf transformation 115] {a:0,w:15,wCS:1500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
143. r2_hidden(k4_tarski(sK1,sK3),sK0) | ~v2_funct_1(sK0) [cnf transformation 115] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
144. r2_hidden(k4_tarski(sK2,sK3),sK0) | ~v2_funct_1(sK0) [cnf transformation 115] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
145. sK1 != sK2 | ~v2_funct_1(sK0) [cnf transformation 115] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
176. k1_funct_1(X0,X3) != k1_funct_1(X0,X4) | X3 = X4 | ~r2_hidden(X4,k9_xtuple_0(X0)) | ~r2_hidden(X3,k9_xtuple_0(X0)) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 119] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
177. r2_hidden(sK4(X0),k9_xtuple_0(X0)) | v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 119] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
178. r2_hidden(sK5(X0),k9_xtuple_0(X0)) | v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 119] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
179. k1_funct_1(X0,sK4(X0)) = k1_funct_1(X0,sK5(X0)) | v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 119] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
180. sK4(X0) != sK5(X0) | v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 119] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
186. ~r2_hidden(k4_tarski(X0,X1),X2) | r2_hidden(X0,k9_xtuple_0(X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [cnf transformation 121] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
187. ~r2_hidden(k4_tarski(X0,X1),X2) | k1_funct_1(X2,X0) = X1 | ~v1_funct_1(X2) | ~v1_relat_1(X2) [cnf transformation 121] {a:0,w:14,wCS:5600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
188. r2_hidden(k4_tarski(X0,X1),X2) | k1_funct_1(X2,X0) != X1 | ~r2_hidden(X0,k9_xtuple_0(X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [cnf transformation 121] {a:0,w:18,wCS:7200,sine:2,thAx:0,allAx:1,thDist:-1}
219. r2_hidden(k4_tarski(X0,k1_funct_1(X2,X0)),X2) | ~r2_hidden(X0,k9_xtuple_0(X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [equality resolution 188] {a:0,w:15,wCS:6000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
221. 1 <=> v2_funct_1(sK0) [avatar definition]
222. v2_funct_1(sK0) <- (1) [avatar component clause 221] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
223. ~v2_funct_1(sK0) <- (~1) [avatar component clause 221] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
225. 2 <=> sK1 = sK2 [avatar definition]
227. sK1 != sK2 <- (~2) [avatar component clause 225] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
228. ~1 | ~2 [avatar split clause 145,225,221]
230. 3 <=> r2_hidden(k4_tarski(sK2,sK3),sK0) [avatar definition]
232. r2_hidden(k4_tarski(sK2,sK3),sK0) <- (3) [avatar component clause 230] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
233. ~1 | 3 [avatar split clause 144,230,221]
235. 4 <=> r2_hidden(k4_tarski(sK1,sK3),sK0) [avatar definition]
237. r2_hidden(k4_tarski(sK1,sK3),sK0) <- (4) [avatar component clause 235] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
238. ~1 | 4 [avatar split clause 143,235,221]
240. 5 <=> ! [X5,X4,X6] : (X4 = X5 | ~r2_hidden(k4_tarski(X4,X6),sK0) | ~r2_hidden(k4_tarski(X5,X6),sK0)) [avatar definition]
241. ~r2_hidden(k4_tarski(X5,X6),sK0) | ~r2_hidden(k4_tarski(X4,X6),sK0) | X4 = X5 <- (5) [avatar component clause 240] {a:0,w:13,wCS:1300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
242. 1 | 5 [avatar split clause 142,240,221]
244. 6 <=> v1_funct_1(sK0) [avatar definition]
246. v1_funct_1(sK0) <- (6) [avatar component clause 244] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
247. 6 [avatar split clause 141,244]
249. 7 <=> v1_relat_1(sK0) [avatar definition]
251. v1_relat_1(sK0) <- (7) [avatar component clause 249] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
252. 7 [avatar split clause 140,249]
476. ~r2_hidden(X0,k9_xtuple_0(sK0)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | ~r2_hidden(k4_tarski(X1,k1_funct_1(sK0,X0)),sK0) | X0 = X1 <- (5) [resolution 219,241] {a:1,w:18,wCS:1800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
486. ~r2_hidden(X0,k9_xtuple_0(sK0)) | ~v1_relat_1(sK0) | ~r2_hidden(k4_tarski(X1,k1_funct_1(sK0,X0)),sK0) | X0 = X1 <- (5, 6) [subsumption resolution 476,246] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
487. ~r2_hidden(k4_tarski(X1,k1_funct_1(sK0,X0)),sK0) | ~r2_hidden(X0,k9_xtuple_0(sK0)) | X0 = X1 <- (5, 6, 7) [subsumption resolution 486,251] {a:1,w:14,wCS:1400,nSel:1,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
492. ~r2_hidden(k4_tarski(X0,k1_funct_1(sK0,sK4(sK0))),sK0) | ~r2_hidden(sK5(sK0),k9_xtuple_0(sK0)) | sK5(sK0) = X0 | v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (5, 6, 7) [superposition 487,179] {a:2,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
498. 39 <=> r2_hidden(sK5(sK0),k9_xtuple_0(sK0)) [avatar definition]
500. ~r2_hidden(sK5(sK0),k9_xtuple_0(sK0)) <- (~39) [avatar component clause 498] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
502. 40 <=> ! [X0] : (~r2_hidden(k4_tarski(X0,k1_funct_1(sK0,sK4(sK0))),sK0) | sK5(sK0) = X0) [avatar definition]
503. ~r2_hidden(k4_tarski(X0,k1_funct_1(sK0,sK4(sK0))),sK0) | sK5(sK0) = X0 <- (40) [avatar component clause 502] {a:2,w:12,wCS:1200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
507. v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~39) [resolution 500,178] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
509. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, ~39) [subsumption resolution 507,223] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
510. ~v1_relat_1(sK0) <- (~1, 6, ~39) [subsumption resolution 509,246] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
511. $false <- (~1, 6, 7, ~39) [subsumption resolution 510,251] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
512. 1 | ~6 | ~7 | 39 [avatar contradiction clause 511]
525. sK5(sK0) = sK4(sK0) | ~r2_hidden(sK4(sK0),k9_xtuple_0(sK0)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (40) [resolution 503,219] {a:3,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
526. sK5(sK0) = sK4(sK0) | ~r2_hidden(sK4(sK0),k9_xtuple_0(sK0)) | ~v1_relat_1(sK0) <- (6, 40) [subsumption resolution 525,246] {a:3,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
527. sK5(sK0) = sK4(sK0) | ~r2_hidden(sK4(sK0),k9_xtuple_0(sK0)) <- (6, 7, 40) [subsumption resolution 526,251] {a:3,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
529. 41 <=> r2_hidden(sK4(sK0),k9_xtuple_0(sK0)) [avatar definition]
531. ~r2_hidden(sK4(sK0),k9_xtuple_0(sK0)) <- (~41) [avatar component clause 529] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
533. 42 <=> sK5(sK0) = sK4(sK0) [avatar definition]
535. sK5(sK0) = sK4(sK0) <- (42) [avatar component clause 533] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
536. ~41 | 42 | ~6 | ~7 | ~40 [avatar split clause 527,502,249,244,533,529]
548. v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~41) [resolution 531,177] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
550. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, ~41) [subsumption resolution 548,223] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
551. ~v1_relat_1(sK0) <- (~1, 6, ~41) [subsumption resolution 550,246] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
552. $false <- (~1, 6, 7, ~41) [subsumption resolution 551,251] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
553. 1 | ~6 | ~7 | 41 [avatar contradiction clause 552]
564. sK4(sK0) != sK4(sK0) | v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (42) [superposition 180,535] {a:4,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
566. v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (42) [trivial inequality removal 564] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
568. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 42) [subsumption resolution 566,223] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
569. ~v1_relat_1(sK0) <- (~1, 6, 42) [subsumption resolution 568,246] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
570. $false <- (~1, 6, 7, 42) [subsumption resolution 569,251] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
571. 1 | ~6 | ~7 | ~42 [avatar contradiction clause 570]
572. ~r2_hidden(k4_tarski(X0,k1_funct_1(sK0,sK4(sK0))),sK0) | ~r2_hidden(sK5(sK0),k9_xtuple_0(sK0)) | sK5(sK0) = X0 | v2_funct_1(sK0) | ~v1_relat_1(sK0) <- (5, 6, 7) [subsumption resolution 492,246] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
573. ~r2_hidden(k4_tarski(X0,k1_funct_1(sK0,sK4(sK0))),sK0) | ~r2_hidden(sK5(sK0),k9_xtuple_0(sK0)) | sK5(sK0) = X0 | v2_funct_1(sK0) <- (5, 6, 7) [subsumption resolution 572,251] {a:2,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
574. 1 | ~39 | 40 | ~5 | ~6 | ~7 [avatar split clause 573,249,244,240,502,498,221]
576. sK3 = k1_funct_1(sK0,sK2) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3) [resolution 232,187] {a:1,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
577. r2_hidden(sK2,k9_xtuple_0(sK0)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (3) [resolution 232,186] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
581. sK3 = k1_funct_1(sK0,sK2) | ~v1_relat_1(sK0) <- (3, 6) [subsumption resolution 576,246] {a:1,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
582. sK3 = k1_funct_1(sK0,sK2) <- (3, 6, 7) [subsumption resolution 581,251] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
584. 43 <=> sK3 = k1_funct_1(sK0,sK2) [avatar definition]
586. sK3 = k1_funct_1(sK0,sK2) <- (43) [avatar component clause 584] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
587. 43 | ~3 | ~6 | ~7 [avatar split clause 582,249,244,230,584]
588. r2_hidden(sK2,k9_xtuple_0(sK0)) | ~v1_relat_1(sK0) <- (3, 6) [subsumption resolution 577,246] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
589. r2_hidden(sK2,k9_xtuple_0(sK0)) <- (3, 6, 7) [subsumption resolution 588,251] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
591. 44 <=> r2_hidden(sK2,k9_xtuple_0(sK0)) [avatar definition]
593. r2_hidden(sK2,k9_xtuple_0(sK0)) <- (44) [avatar component clause 591] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
594. 44 | ~3 | ~6 | ~7 [avatar split clause 589,249,244,230,591]
613. sK3 = k1_funct_1(sK0,sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (4) [resolution 237,187] {a:1,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
614. r2_hidden(sK1,k9_xtuple_0(sK0)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (4) [resolution 237,186] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
620. sK3 = k1_funct_1(sK0,sK1) | ~v1_relat_1(sK0) <- (4, 6) [subsumption resolution 613,246] {a:1,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
621. sK3 = k1_funct_1(sK0,sK1) <- (4, 6, 7) [subsumption resolution 620,251] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
622. r2_hidden(sK1,k9_xtuple_0(sK0)) | ~v1_relat_1(sK0) <- (4, 6) [subsumption resolution 614,246] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
623. r2_hidden(sK1,k9_xtuple_0(sK0)) <- (4, 6, 7) [subsumption resolution 622,251] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
636. 50 <=> r2_hidden(sK1,k9_xtuple_0(sK0)) [avatar definition]
638. r2_hidden(sK1,k9_xtuple_0(sK0)) <- (50) [avatar component clause 636] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
639. 50 | ~4 | ~6 | ~7 [avatar split clause 623,249,244,235,636]
641. 51 <=> sK3 = k1_funct_1(sK0,sK1) [avatar definition]
643. sK3 = k1_funct_1(sK0,sK1) <- (51) [avatar component clause 641] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
644. 51 | ~4 | ~6 | ~7 [avatar split clause 621,249,244,235,641]
679. sK3 != k1_funct_1(sK0,X0) | sK2 = X0 | ~r2_hidden(sK2,k9_xtuple_0(sK0)) | ~r2_hidden(X0,k9_xtuple_0(sK0)) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (43) [superposition 176,586] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
683. sK3 != k1_funct_1(sK0,X0) | sK2 = X0 | ~r2_hidden(X0,k9_xtuple_0(sK0)) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (43, 44) [subsumption resolution 679,593] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
684. sK3 != k1_funct_1(sK0,X0) | sK2 = X0 | ~r2_hidden(X0,k9_xtuple_0(sK0)) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (1, 43, 44) [subsumption resolution 683,222] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
685. sK3 != k1_funct_1(sK0,X0) | sK2 = X0 | ~r2_hidden(X0,k9_xtuple_0(sK0)) | ~v1_relat_1(sK0) <- (1, 6, 43, 44) [subsumption resolution 684,246] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
686. sK3 != k1_funct_1(sK0,X0) | sK2 = X0 | ~r2_hidden(X0,k9_xtuple_0(sK0)) <- (1, 6, 7, 43, 44) [subsumption resolution 685,251] {a:2,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
798. sK3 != sK3 | sK1 = sK2 | ~r2_hidden(sK1,k9_xtuple_0(sK0)) <- (1, 6, 7, 43, 44, 51) [superposition 686,643] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
800. sK1 = sK2 | ~r2_hidden(sK1,k9_xtuple_0(sK0)) <- (1, 6, 7, 43, 44, 51) [trivial inequality removal 798] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
802. ~r2_hidden(sK1,k9_xtuple_0(sK0)) <- (1, ~2, 6, 7, 43, 44, 51) [subsumption resolution 800,227] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
803. $false <- (1, ~2, 6, 7, 43, 44, 50, 51) [subsumption resolution 802,638] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
804. ~1 | 2 | ~6 | ~7 | ~43 | ~44 | ~50 | ~51 [avatar contradiction clause 803]
805. $false [avatar sat refutation 228,233,238,242,247,252,512,536,553,571,574,587,594,639,644,804] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:469,thDist:-469}
% SZS output end Proof for grfunc_1__t9_grfunc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 53
% >>> Preprocessing
% Introduced skolems: 15
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 79
% Generated clauses: 391
% Active clauses: 120
% Passive clauses: 167
% Final active clauses: 112
% Final passive clauses: 34
% Main loop iterations started: 132
% 
% >>> Simplifying Inferences
% Duplicate literals: 23
% Trivial inequalities: 2
% Fw subsumption resolutions: 63
% Fw demodulations: 2
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Equational tautologies: 3
% Forward subsumptions: 56
% 
% >>> Generating Inferences
% Binary resolution: 88
% Unit resulting resolution: 6
% Forward superposition: 41
% Backward superposition: 23
% Self superposition: 4
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 15
% Split components: 31
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 75
% SAT solver unit clauses: 25
% SAT solver binary clauses: 22
% 
% Memory used [KB]: 5373
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.004 s ( own 0.003 s ) 

