i: [1,33,0,7,2,1,0,0] t18_gate_1
a: [3,46,0,2,1,33]
s: 46
i: [1,32,0,7,2,1,0,0] t18_gate_1
a: [3,55,0,2,1,32]
s: 55
i: [1,31,0,7,2,1,0,0] t18_gate_1
a: [3,60,0,2,1,31]
s: 60
a: [3,49,0,5,1,33]
s: 49
i: [1,34,0,2,1,0,0,255] fc1_xboole_0
a: [3,66,0,2,1,34]
s: 66
i: [1,41,0,2,1,0,0,1] rc2_xboole_0
a: [3,71,0,2,1,41]
s: 71
i: [1,42,0,2,1,0,0,1] rc1_xboole_0
a: [3,76,0,2,1,42]
s: 76
i: [1,35,0,5,2,0,0,255] fc2_gate_1
s: 35
i: [1,36,0,5,2,0,0,255] fc1_gate_1
s: 36
i: [1,37,0,9,2,0,0,255] d12_gate_1
s: 37
i: [1,38,0,9,2,0,0,255] d12_gate_1
s: 38
i: [1,39,0,9,2,0,0,255] d12_gate_1
s: 39
i: [1,40,0,12,4,0,0,255] d12_gate_1
s: 40
d: [2,85,1,8,4,1,44,49,40]
d: [2,89,1,6,3,2,21,85,66]
d: [2,90,1,4,2,3,21,89,46]
d: [2,91,1,2,1,4,21,90,55]
d: [2,92,1,0,0,5,21,91,60]
e: 92
r: 49
r: 55
a: [3,54,0,2,1,32]
s: 54
a: [3,50,0,5,1,33]
s: 50
d: [2,96,1,5,2,1,44,50,38]
d: [2,98,1,3,1,2,21,96,54]
a: [3,102,1,3,1,98]
s: 102
d: [2,104,2,0,0,2,52,66,102,36]
e: 104
d: [2,105,2,2,1,1,40,102,36]
d: [2,107,2,0,0,2,21,105,66]
e: 107
r: 54
r: 60
r: 102
r: 46
i: [1,26,0,11,4,1,0,0] t18_gate_1
d: [2,97,1,5,2,1,44,50,37]
d: [2,95,1,5,2,1,44,50,39]
f: 95,97,96,107,92,34,26,31,32,33
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t18_gate_1
% SZS output start Proof for gate_1__t18_gate_1
1. ! [X0] : ! [X1] : ! [X2] : (~(v1_xboole_0(k12_gate_1(X0,X1,X2)) & ~(~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & ~(~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k12_gate_1(X0,X1,X2)))) [input t18_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : (~(v1_xboole_0(k12_gate_1(X0,X1,X2)) & ~(~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & ~(~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k12_gate_1(X0,X1,X2)))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : (~(k12_gate_1(X0,X1,X2) != k1_xboole_0 & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & (~(~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k12_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0))) [input d12_gate_1]
14. ~! [X0,X1,X2] : (~(v1_xboole_0(k12_gate_1(X0,X1,X2)) & ~(~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & ~(~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k12_gate_1(X0,X1,X2)))) [flattening 2]
15. ! [X0,X1,X2] : (~(k12_gate_1(X0,X1,X2) != k1_xboole_0 & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & (~(~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k12_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0))) [flattening 13]
16. ? [X0,X1,X2] : ((v1_xboole_0(k12_gate_1(X0,X1,X2)) & (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) | (~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k12_gate_1(X0,X1,X2)))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2] : ((k12_gate_1(X0,X1,X2) = k1_xboole_0 | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) & (k12_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0) | (~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)))) [ennf transformation 15]
20. ? [X0,X1,X2] : ((v1_xboole_0(k12_gate_1(X0,X1,X2)) & (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) | (~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k12_gate_1(X0,X1,X2)))) => ((v1_xboole_0(k12_gate_1(sK0,sK1,sK2)) & (v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0))) | (~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0) & ~v1_xboole_0(k12_gate_1(sK0,sK1,sK2)))) [choice axiom]
21. (v1_xboole_0(k12_gate_1(sK0,sK1,sK2)) & (v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0))) | (~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0) & ~v1_xboole_0(k12_gate_1(sK0,sK1,sK2))) [skolemisation 16,20]
26. v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | ~v1_xboole_0(k12_gate_1(sK0,sK1,sK2)) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. v1_xboole_0(k12_gate_1(sK0,sK1,sK2)) | ~v1_xboole_0(sK0) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. v1_xboole_0(k12_gate_1(sK0,sK1,sK2)) | ~v1_xboole_0(sK1) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. v1_xboole_0(k12_gate_1(sK0,sK1,sK2)) | ~v1_xboole_0(sK2) [cnf transformation 21] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
36. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
37. k12_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. k12_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X1) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. k12_gate_1(X0,X1,X2) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X2) [cnf transformation 19] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. k12_gate_1(X0,X1,X2) = k1_xboole_0 | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
44. 1 <=> v1_xboole_0(sK2) [avatar definition]
46. ~v1_xboole_0(sK2) <- (~1) [avatar component clause 44] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
48. 2 <=> v1_xboole_0(k12_gate_1(sK0,sK1,sK2)) [avatar definition]
49. ~v1_xboole_0(k12_gate_1(sK0,sK1,sK2)) <- (~2) [avatar component clause 48] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
50. v1_xboole_0(k12_gate_1(sK0,sK1,sK2)) <- (2) [avatar component clause 48] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
51. ~1 | 2 [avatar split clause 33,48,44]
53. 3 <=> v1_xboole_0(sK1) [avatar definition]
55. ~v1_xboole_0(sK1) <- (~3) [avatar component clause 53] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. ~3 | 2 [avatar split clause 32,48,53]
58. 4 <=> v1_xboole_0(sK0) [avatar definition]
60. ~v1_xboole_0(sK0) <- (~4) [avatar component clause 58] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
61. ~4 | 2 [avatar split clause 31,48,58]
62. ~2 | 4 | 3 | 1 [avatar split clause 26,44,53,58,48]
64. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
66. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 64] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
67. 5 [avatar split clause 34,64]
85. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~2) [superposition 49,40] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
89. v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~2, 5) [subsumption resolution 85,66] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
90. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1, ~2, 5) [subsumption resolution 89,46] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. v1_xboole_0(sK0) <- (~1, ~2, ~3, 5) [subsumption resolution 90,55] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. $false <- (~1, ~2, ~3, ~4, 5) [subsumption resolution 91,60] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
93. 1 | 2 | 3 | 4 | ~5 [avatar contradiction clause 92]
95. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK2) <- (2) [superposition 50,39] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK1) <- (2) [superposition 50,38] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
97. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK0) <- (2) [superposition 50,37] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. 8 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
102. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (8) [avatar component clause 100] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. ~v1_xboole_0(k1_xboole_0) <- (8) [resolution 102,36] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
107. $false <- (5, 8) [subsumption resolution 105,66] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
108. ~5 | ~8 [avatar contradiction clause 107]
109. ~3 | 8 | ~2 [avatar split clause 96,48,100,53]
110. ~4 | 8 | ~2 [avatar split clause 97,48,100,58]
111. ~1 | 8 | ~2 [avatar split clause 95,48,100,44]
112. $false [avatar sat refutation 51,56,61,62,67,93,108,109,110,111] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:62,thDist:-62}
% SZS output end Proof for gate_1__t18_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 58
% Active clauses: 16
% Passive clauses: 21
% Final active clauses: 10
% Final passive clauses: 4
% Main loop iterations started: 16
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 7
% Backward superposition: 7
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 16
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 3
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% superposition: 0.001 s ( own 0.0000 s ) 
% other: 0.001 s ( own 0.0000 s ) 

