i: [1,109,0,3,1,1,0,0] t32_setfam_1
a: [3,167,0,3,1,109]
s: 167
i: [1,110,0,5,1,1,0,0] t32_setfam_1
a: [3,162,0,5,1,110]
s: 162
i: [1,108,0,5,1,1,0,0] t32_setfam_1
a: [3,172,0,5,1,108]
s: 172
i: [1,111,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,177,0,2,1,111]
s: 177
i: [1,112,0,2,1,0,0,1] fc1_xboole_0
a: [3,182,0,2,1,112]
s: 182
i: [1,155,0,2,1,0,0,2] rc2_xboole_0
a: [3,192,0,2,1,155]
s: 192
i: [1,156,0,2,1,0,0,2] rc1_xboole_0
a: [3,197,0,2,1,156]
s: 197
i: [1,114,0,3,1,0,0,2] fc1_subset_1
s: 114
i: [1,145,0,3,1,0,0,2] rc2_subset_1
s: 145
i: [1,148,0,3,1,0,0,255] reflexivity_r1_tarski
s: 148
i: [1,113,0,3,1,0,0,255] d2_xboole_0
a: [3,187,0,3,1,113]
s: 187
r: 177
i: [1,143,0,4,1,0,0,1] existence_m1_subset_1
s: 143
i: [1,147,0,4,1,0,0,255] rc3_subset_1
s: 147
i: [1,115,0,5,1,0,0,255] t4_boole
s: 115
i: [1,116,0,5,1,0,0,255] t3_boole
s: 116
i: [1,120,0,5,2,0,0,2] rc1_subset_1
s: 120
i: [1,123,0,5,2,0,0,1] t6_boole
s: 123
i: [1,140,0,5,2,0,0,2] t7_boole
s: 140
i: [1,144,0,5,1,0,0,2] rc2_subset_1
s: 144
i: [1,146,0,5,1,0,0,255] rc3_subset_1
s: 146
i: [1,122,0,6,2,0,0,255] rc4_subset_1
s: 122
i: [1,126,0,6,2,0,0,2] t1_subset
s: 126
i: [1,149,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 149
i: [1,119,0,7,2,0,0,2] rc1_subset_1
s: 119
i: [1,121,0,7,2,0,0,255] rc4_subset_1
s: 121
i: [1,137,0,7,2,0,0,255] t3_subset
s: 137
d: [2,204,1,4,1,1,40,137,172]
a: [3,213,1,4,1,204]
s: 213
i: [1,138,0,7,2,0,0,255] t3_subset
s: 138
d: [2,208,1,4,1,0,40,137,144]
s: 208
i: [1,139,0,7,3,0,0,1] t8_boole
s: 139
i: [1,124,0,8,3,0,0,2] cc1_subset_1
s: 124
d: [2,209,1,4,1,0,40,137,146]
s: 209
i: [1,127,0,8,3,0,0,2] t2_subset
s: 127
d: [2,224,1,9,2,1,40,127,172]
d: [2,231,1,5,1,1,21,224,114]
a: [3,235,1,5,1,231]
s: 235
i: [1,135,0,8,2,0,0,255] d3_tarski
s: 135
i: [1,136,0,8,2,0,0,255] d3_tarski
s: 136
d: [2,241,2,5,1,1,40,235,149]
a: [3,247,2,5,1,241]
s: 247
i: [1,125,0,9,3,0,0,255] cc4_subset_1
s: 125
d: [2,207,1,5,1,0,40,137,143]
s: 207
i: [1,134,0,9,3,0,0,255] d3_tarski
s: 134
i: [1,142,0,9,3,0,0,2] t5_subset
s: 142
d: [2,262,2,7,2,1,40,134,213]
s: 262
i: [1,141,0,10,3,0,0,2] t4_subset
s: 141
d: [2,278,1,7,2,1,40,141,172]
s: 278
i: [1,150,0,10,2,0,0,2] dt_k3_subset_1
s: 150
i: [1,152,0,10,2,0,0,255] dt_o_2_1_setfam_1
s: 152
d: [2,285,2,8,3,1,40,278,142]
a: [3,293,2,2,1,285]
s: 293
i: [1,117,0,11,4,0,0,255] cc3_subset_1
s: 117
d: [2,331,1,9,3,1,40,117,172]
d: [2,340,1,6,2,1,21,331,114]
a: [3,348,1,2,1,340]
s: 348
i: [1,118,0,11,4,0,0,255] cc2_subset_1
s: 118
i: [1,128,0,11,2,0,0,255] d4_subset_1
s: 128
d: [2,315,1,5,1,1,40,152,172]
a: [3,324,1,5,1,315]
s: 324
r: 348
i: [1,151,0,11,2,0,0,2] involutiveness_k3_subset_1
s: 151
a: [3,347,1,2,1,340]
s: 347
d: [2,365,2,0,0,2,52,167,347,123]
e: 365
d: [2,366,2,0,0,3,52,182,167,347,139]
e: 366
d: [2,367,2,0,0,3,52,182,167,347,139]
e: 367
r: 347
i: [1,153,0,12,2,0,0,1] involutiveness_k7_setfam_1
s: 153
i: [1,154,0,12,2,0,0,1] dt_k7_setfam_1
s: 154
s: 348
i: [1,131,0,22,4,0,0,1] d7_setfam_1
s: 131
a: [3,344,1,4,1,340]
s: 344
i: [1,157,0,26,5,0,0,1] d7_setfam_1
s: 157
d: [2,405,0,19,4,0,21,157,154]
s: 405
d: [2,371,1,10,2,1,44,153,162]
d: [2,372,1,5,1,2,21,371,172]
a: [3,376,1,5,1,372]
s: 376
i: [1,158,0,26,5,0,0,1] d7_setfam_1
s: 158
d: [2,386,1,10,2,1,44,154,162]
d: [2,390,1,5,1,2,21,386,172]
a: [3,394,1,5,1,390]
s: 394
d: [2,412,0,22,4,0,21,158,141]
d: [2,413,0,15,3,0,21,412,154]
s: 413
i: [1,132,0,29,5,0,0,1] d7_setfam_1
s: 132
d: [2,415,2,9,3,1,40,394,117]
d: [2,428,2,7,2,2,21,415,182]
d: [2,429,2,4,1,2,21,428,114]
a: [3,433,2,4,1,429]
s: 433
i: [1,133,0,29,5,0,0,1] d7_setfam_1
s: 133
d: [2,420,2,4,1,1,40,394,137]
a: [3,438,2,4,1,420]
s: 438
d: [2,447,1,13,3,1,44,413,162]
d: [2,450,1,8,2,2,21,447,172]
s: 450
d: [2,222,1,5,2,0,40,124,146]
s: 222
d: [2,356,2,7,2,1,40,324,127]
d: [2,357,2,5,1,2,21,356,348]
a: [3,361,2,5,1,357]
s: 361
d: [2,228,1,8,2,0,40,127,144]
d: [2,239,1,5,1,0,21,228,114]
s: 239
d: [2,414,2,5,1,1,40,394,152]
a: [3,426,2,5,1,414]
s: 426
d: [2,229,1,8,2,0,40,127,146]
d: [2,240,1,5,1,0,21,229,114]
s: 240
d: [2,251,1,5,2,0,40,135,140]
s: 251
d: [2,421,2,9,2,1,40,394,127]
d: [2,440,2,5,1,1,21,421,114]
a: [3,444,2,5,1,440]
s: 444
d: [2,336,1,9,3,0,40,117,146]
d: [2,351,1,5,2,0,21,336,147]
s: 351
d: [2,476,3,5,1,1,40,361,149]
a: [3,482,3,5,1,476]
s: 482
d: [2,205,1,6,2,0,40,137,119]
s: 205
d: [2,206,1,6,2,0,40,137,121]
s: 206
d: [2,493,3,5,1,1,40,444,149]
a: [3,500,3,5,1,493]
s: 500
d: [2,220,1,6,2,0,40,124,143]
s: 220
d: [2,288,2,6,2,1,40,278,137]
s: 288
d: [2,227,1,6,2,0,40,127,143]
s: 227
d: [2,259,1,6,2,0,40,125,144]
s: 259
d: [2,519,3,6,2,1,40,227,288]
d: [2,520,3,4,1,2,21,519,348]
a: [3,524,3,4,1,520]
s: 524
d: [2,271,1,6,2,0,40,142,144]
s: 271
d: [2,508,3,5,1,2,40,288,361]
a: [3,513,3,5,1,508]
s: 513
d: [2,272,1,6,2,0,40,142,146]
s: 272
d: [2,335,1,9,3,0,40,117,144]
d: [2,350,1,6,2,0,21,335,145]
s: 350
d: [2,460,3,6,2,1,44,433,123]
s: 460
d: [2,200,1,7,3,0,44,123,123]
s: 200
d: [2,535,4,4,1,1,40,460,147]
a: [3,540,4,4,1,535]
s: 540
d: [2,201,1,7,2,0,44,116,123]
s: 201
d: [2,202,1,7,2,0,44,115,123]
s: 202
d: [2,465,3,6,2,1,44,438,123]
s: 465
d: [2,216,1,7,3,0,40,124,138]
s: 216
d: [2,409,2,7,2,1,44,376,123]
s: 409
d: [2,223,1,10,3,0,40,127,138]
d: [2,230,1,7,2,0,21,223,114]
s: 230
d: [2,225,1,10,3,0,40,127,119]
d: [2,237,1,7,2,0,21,225,114]
s: 237
d: [2,416,2,7,2,1,40,394,141]
s: 416
d: [2,226,1,10,3,0,40,127,121]
d: [2,238,1,7,2,0,21,226,114]
s: 238
d: [2,588,3,6,2,1,40,416,137]
s: 588
d: [2,258,1,7,2,0,40,125,143]
s: 258
d: [2,270,1,7,2,0,40,142,143]
s: 270
d: [2,422,2,7,2,1,44,394,123]
s: 422
d: [2,282,1,7,2,0,40,141,144]
s: 282
d: [2,273,3,7,2,1,40,262,149]
s: 273
d: [2,283,1,7,2,0,40,141,146]
s: 283
d: [2,249,1,8,2,0,40,135,149]
s: 249
d: [2,616,4,8,2,1,40,273,262]
d: [2,618,4,4,1,1,18,616]
a: [3,622,4,4,1,618]
s: 622
d: [2,250,1,8,2,0,40,135,126]
s: 250
d: [2,274,3,7,2,1,40,262,126]
s: 274
d: [2,254,1,8,3,0,40,125,138]
s: 254
d: [2,266,1,8,3,0,40,142,138]
s: 266
d: [2,464,3,7,2,1,40,438,134]
s: 464
d: [2,281,1,8,2,0,40,141,143]
s: 281
d: [2,488,3,7,2,1,44,426,123]
s: 488
d: [2,277,1,9,3,0,40,141,138]
s: 277
d: [2,279,1,9,3,0,40,141,119]
s: 279
d: [2,496,3,7,2,1,44,444,123]
s: 496
d: [2,280,1,9,3,0,40,141,121]
s: 280
d: [2,561,3,15,4,1,44,413,409]
d: [2,565,3,10,3,1,21,561,140]
a: [3,568,3,7,2,565]
s: 568
d: [2,310,1,9,2,0,40,150,137]
s: 310
d: [2,313,1,9,2,0,40,152,138]
s: 313
d: [2,700,3,2,1,2,21,568,422]
s: 700
r: 192
r: 540
r: 351
r: 120
r: 348
r: 293
r: 114
r: 496
r: 488
r: 460
r: 422
r: 409
r: 272
r: 271
r: 270
r: 266
r: 259
r: 258
r: 254
r: 251
r: 222
r: 220
r: 216
r: 202
r: 201
r: 142
r: 140
r: 139
r: 125
r: 124
r: 123
r: 117
d: [2,718,4,0,0,3,40,700,197]
e: 718
d: [2,717,4,0,0,2,40,700,145]
e: 717
d: [2,716,4,0,0,3,40,700,182]
e: 716
d: [2,715,4,0,0,3,52,182,700]
e: 715
r: 700
s: 192
s: 348
s: 114
s: 120
s: 293
s: 123
a: [3,571,3,3,1,565]
s: 571
r: 262
r: 278
r: 288
r: 273
r: 622
r: 247
d: [2,741,4,0,0,2,40,571,361]
e: 741
d: [2,735,4,0,0,2,52,361,571]
e: 735
d: [2,736,4,0,0,2,52,361,148,571,134]
e: 736
d: [2,737,4,0,0,3,52,324,348,571,127]
e: 737
d: [2,738,4,0,0,2,52,348,571,227]
e: 738
d: [2,743,4,2,1,1,40,571,227]
d: [2,749,4,0,0,2,21,743,348]
e: 749
f: 749,717,565,390,372,367,112,108,109,110
% Refutation found. Thanks to Tanya!
% SZS status Theorem for setfam_1__t32_setfam_1
% SZS output start Proof for setfam_1__t32_setfam_1
1. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => ~(k1_xboole_0 = k7_setfam_1(X0,X1) & k1_xboole_0 != X1)) [input t32_setfam_1]
2. ~! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => ~(k1_xboole_0 = k7_setfam_1(X0,X1) & k1_xboole_0 != X1)) [negated conjecture 1]
11. ! [X0] : ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) => (k7_setfam_1(X0,X1) = X2 <=> ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(X0)) => (r2_hidden(X3,X2) <=> r2_hidden(k3_subset_1(X0,X3),X1)))))) [input d7_setfam_1]
16. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => m1_subset_1(k7_setfam_1(X0,X1),k1_zfmisc_1(k1_zfmisc_1(X0)))) [input dt_k7_setfam_1]
20. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
22. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
24. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => k7_setfam_1(X0,k7_setfam_1(X0,X1)) = X1) [input involutiveness_k7_setfam_1]
27. ! [X0] : ? [X1] : (v1_xboole_0(X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) [input rc2_subset_1]
33. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
37. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
39. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
40. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
41. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [input t8_boole]
42. ~! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => ~(k1_xboole_0 = k7_setfam_1(X0,X1) & k1_xboole_0 != X1)) [flattening 2]
44. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 33]
46. ! [X0,X1] : (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) => ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) => (k7_setfam_1(X0,X1) = X2 <=> ! [X3] : (m1_subset_1(X3,k1_zfmisc_1(X0)) => (r2_hidden(X3,X2) <=> r2_hidden(k3_subset_1(X0,X3),X1)))))) [flattening 11]
49. ! [X0,X1] : ~(v1_xboole_0(X1) & X0 != X1 & v1_xboole_0(X0)) [flattening 41]
50. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 40]
51. ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 37]
54. ? [X0,X1] : ((k1_xboole_0 = k7_setfam_1(X0,X1) & k1_xboole_0 != X1) & m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [ennf transformation 42]
55. ? [X0,X1] : (k1_xboole_0 = k7_setfam_1(X0,X1) & k1_xboole_0 != X1 & m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [flattening 54]
62. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 39]
66. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 44]
67. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 66]
69. ! [X0,X1] : (! [X2] : ((k7_setfam_1(X0,X1) = X2 <=> ! [X3] : ((r2_hidden(X3,X2) <=> r2_hidden(k3_subset_1(X0,X3),X1)) | ~m1_subset_1(X3,k1_zfmisc_1(X0)))) | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0)))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [ennf transformation 46]
71. ! [X0,X1] : (~v1_xboole_0(X1) | X0 = X1 | ~v1_xboole_0(X0)) [ennf transformation 49]
72. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 50]
73. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))) [ennf transformation 51]
74. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [flattening 73]
80. ! [X0,X1] : (k7_setfam_1(X0,k7_setfam_1(X0,X1)) = X1 | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [ennf transformation 24]
81. ! [X0,X1] : (m1_subset_1(k7_setfam_1(X0,X1),k1_zfmisc_1(k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [ennf transformation 16]
82. ? [X0,X1] : (k1_xboole_0 = k7_setfam_1(X0,X1) & k1_xboole_0 != X1 & m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) => (k1_xboole_0 = k7_setfam_1(sK0,sK1) & k1_xboole_0 != sK1 & m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0)))) [choice axiom]
83. k1_xboole_0 = k7_setfam_1(sK0,sK1) & k1_xboole_0 != sK1 & m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) [skolemisation 55,82]
88. ! [X0,X1] : (! [X2] : (((k7_setfam_1(X0,X1) = X2 | ? [X3] : (((~r2_hidden(k3_subset_1(X0,X3),X1) | ~r2_hidden(X3,X2)) & (r2_hidden(k3_subset_1(X0,X3),X1) | r2_hidden(X3,X2))) & m1_subset_1(X3,k1_zfmisc_1(X0)))) & (! [X3] : (((r2_hidden(X3,X2) | ~r2_hidden(k3_subset_1(X0,X3),X1)) & (r2_hidden(k3_subset_1(X0,X3),X1) | ~r2_hidden(X3,X2))) | ~m1_subset_1(X3,k1_zfmisc_1(X0))) | k7_setfam_1(X0,X1) != X2)) | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0)))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [nnf transformation 69]
89. ! [X0,X1] : (! [X2] : (((k7_setfam_1(X0,X1) = X2 | ? [X3] : ((~r2_hidden(k3_subset_1(X0,X3),X1) | ~r2_hidden(X3,X2)) & (r2_hidden(k3_subset_1(X0,X3),X1) | r2_hidden(X3,X2)) & m1_subset_1(X3,k1_zfmisc_1(X0)))) & (! [X3] : (((r2_hidden(X3,X2) | ~r2_hidden(k3_subset_1(X0,X3),X1)) & (r2_hidden(k3_subset_1(X0,X3),X1) | ~r2_hidden(X3,X2))) | ~m1_subset_1(X3,k1_zfmisc_1(X0))) | k7_setfam_1(X0,X1) != X2)) | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0)))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [flattening 88]
90. ! [X0,X1] : (! [X2] : (((k7_setfam_1(X0,X1) = X2 | ? [X3] : ((~r2_hidden(k3_subset_1(X0,X3),X1) | ~r2_hidden(X3,X2)) & (r2_hidden(k3_subset_1(X0,X3),X1) | r2_hidden(X3,X2)) & m1_subset_1(X3,k1_zfmisc_1(X0)))) & (! [X4] : (((r2_hidden(X4,X2) | ~r2_hidden(k3_subset_1(X0,X4),X1)) & (r2_hidden(k3_subset_1(X0,X4),X1) | ~r2_hidden(X4,X2))) | ~m1_subset_1(X4,k1_zfmisc_1(X0))) | k7_setfam_1(X0,X1) != X2)) | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0)))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [rectify 89]
91. ! [X2,X1,X0] : (? [X3] : ((~r2_hidden(k3_subset_1(X0,X3),X1) | ~r2_hidden(X3,X2)) & (r2_hidden(k3_subset_1(X0,X3),X1) | r2_hidden(X3,X2)) & m1_subset_1(X3,k1_zfmisc_1(X0))) => ((~r2_hidden(k3_subset_1(X0,sK4(X0,X1,X2)),X1) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (r2_hidden(k3_subset_1(X0,sK4(X0,X1,X2)),X1) | r2_hidden(sK4(X0,X1,X2),X2)) & m1_subset_1(sK4(X0,X1,X2),k1_zfmisc_1(X0)))) [choice axiom]
92. ! [X0,X1] : (! [X2] : (((k7_setfam_1(X0,X1) = X2 | ((~r2_hidden(k3_subset_1(X0,sK4(X0,X1,X2)),X1) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (r2_hidden(k3_subset_1(X0,sK4(X0,X1,X2)),X1) | r2_hidden(sK4(X0,X1,X2),X2)) & m1_subset_1(sK4(X0,X1,X2),k1_zfmisc_1(X0)))) & (! [X4] : (((r2_hidden(X4,X2) | ~r2_hidden(k3_subset_1(X0,X4),X1)) & (r2_hidden(k3_subset_1(X0,X4),X1) | ~r2_hidden(X4,X2))) | ~m1_subset_1(X4,k1_zfmisc_1(X0))) | k7_setfam_1(X0,X1) != X2)) | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0)))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0)))) [skolemisation 90,91]
98. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK6(X0),X0)) [choice axiom]
99. ! [X0] : m1_subset_1(sK6(X0),X0) [skolemisation 20,98]
100. ! [X0] : (? [X1] : (v1_xboole_0(X1) & m1_subset_1(X1,k1_zfmisc_1(X0))) => (v1_xboole_0(sK7(X0)) & m1_subset_1(sK7(X0),k1_zfmisc_1(X0)))) [choice axiom]
101. ! [X0] : (v1_xboole_0(sK7(X0)) & m1_subset_1(sK7(X0),k1_zfmisc_1(X0))) [skolemisation 27,100]
108. m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) [cnf transformation 83] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
109. k1_xboole_0 != sK1 [cnf transformation 83] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
110. k1_xboole_0 = k7_setfam_1(sK0,sK1) [cnf transformation 83] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
112. v1_xboole_0(k1_xboole_0) [cnf transformation 22] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
123. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 62] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
127. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 67] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
129. r2_hidden(k3_subset_1(X0,X4),X1) | ~r2_hidden(X4,X2) | ~m1_subset_1(X4,k1_zfmisc_1(X0)) | k7_setfam_1(X0,X1) != X2 | ~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) [cnf transformation 92] {a:0,w:27,wCS:10800,sine:1,thAx:0,allAx:2,thDist:-2}
139. X0 = X1 | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 71] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
140. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 72] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
141. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1) [cnf transformation 74] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
143. m1_subset_1(sK6(X0),X0) [cnf transformation 99] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
145. v1_xboole_0(sK7(X0)) [cnf transformation 101] {a:0,w:3,wCS:1200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
153. k7_setfam_1(X0,k7_setfam_1(X0,X1)) = X1 | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) [cnf transformation 80] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
154. m1_subset_1(k7_setfam_1(X0,X1),k1_zfmisc_1(k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) [cnf transformation 81] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
158. r2_hidden(k3_subset_1(X0,X4),X1) | ~r2_hidden(X4,k7_setfam_1(X0,X1)) | ~m1_subset_1(X4,k1_zfmisc_1(X0)) | ~m1_subset_1(k7_setfam_1(X0,X1),k1_zfmisc_1(k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) [equality resolution 129] {a:0,w:26,wCS:10400,sine:1,thAx:0,allAx:2,thDist:-2}
160. 1 <=> k1_xboole_0 = k7_setfam_1(sK0,sK1) [avatar definition]
162. k1_xboole_0 = k7_setfam_1(sK0,sK1) <- (1) [avatar component clause 160] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
163. 1 [avatar split clause 110,160]
165. 2 <=> k1_xboole_0 = sK1 [avatar definition]
167. k1_xboole_0 != sK1 <- (~2) [avatar component clause 165] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
168. ~2 [avatar split clause 109,165]
170. 3 <=> m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) [avatar definition]
172. m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) <- (3) [avatar component clause 170] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
173. 3 [avatar split clause 108,170]
180. 5 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
182. v1_xboole_0(k1_xboole_0) <- (5) [avatar component clause 180] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
183. 5 [avatar split clause 112,180]
227. r2_hidden(sK6(X4),X4) | v1_xboole_0(X4) [resolution 127,143] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
346. 18 <=> v1_xboole_0(sK1) [avatar definition]
347. v1_xboole_0(sK1) <- (18) [avatar component clause 346] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
348. ~v1_xboole_0(sK1) <- (~18) [avatar component clause 346] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
367. $false <- (~2, 5, 18) [unit resulting resolution 182,167,347,139] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
370. 2 | ~5 | ~18 [avatar contradiction clause 367]
371. sK1 = k7_setfam_1(sK0,k1_xboole_0) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) <- (1) [superposition 153,162] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
372. sK1 = k7_setfam_1(sK0,k1_xboole_0) <- (1, 3) [subsumption resolution 371,172] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
374. 20 <=> sK1 = k7_setfam_1(sK0,k1_xboole_0) [avatar definition]
376. sK1 = k7_setfam_1(sK0,k1_xboole_0) <- (20) [avatar component clause 374] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
377. 20 | ~1 | ~3 [avatar split clause 372,170,160,374]
386. m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_zfmisc_1(sK0))) | ~m1_subset_1(sK1,k1_zfmisc_1(k1_zfmisc_1(sK0))) <- (1) [superposition 154,162] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
390. m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_zfmisc_1(sK0))) <- (1, 3) [subsumption resolution 386,172] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
392. 21 <=> m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_zfmisc_1(sK0))) [avatar definition]
394. m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_zfmisc_1(sK0))) <- (21) [avatar component clause 392] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
395. 21 | ~1 | ~3 [avatar split clause 390,170,160,392]
409. sK1 = k7_setfam_1(sK0,X0) | ~v1_xboole_0(X0) <- (20) [superposition 376,123] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
412. r2_hidden(k3_subset_1(X0,X4),X1) | ~r2_hidden(X4,k7_setfam_1(X0,X1)) | ~m1_subset_1(k7_setfam_1(X0,X1),k1_zfmisc_1(k1_zfmisc_1(X0))) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) [subsumption resolution 158,141] {a:0,w:22,wCS:8800,sine:1,thAx:0,allAx:2,thDist:-2}
413. ~r2_hidden(X4,k7_setfam_1(X0,X1)) | r2_hidden(k3_subset_1(X0,X4),X1) | ~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X0))) [subsumption resolution 412,154] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
422. m1_subset_1(X0,k1_zfmisc_1(k1_zfmisc_1(sK0))) | ~v1_xboole_0(X0) <- (21) [superposition 394,123] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
561. ~r2_hidden(X1,sK1) | r2_hidden(k3_subset_1(sK0,X1),X0) | ~m1_subset_1(X0,k1_zfmisc_1(k1_zfmisc_1(sK0))) | ~v1_xboole_0(X0) <- (20) [superposition 413,409] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
565. ~r2_hidden(X1,sK1) | ~m1_subset_1(X0,k1_zfmisc_1(k1_zfmisc_1(sK0))) | ~v1_xboole_0(X0) <- (20) [subsumption resolution 561,140] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
567. 31 <=> ! [X0] : (~m1_subset_1(X0,k1_zfmisc_1(k1_zfmisc_1(sK0))) | ~v1_xboole_0(X0)) [avatar definition]
568. ~m1_subset_1(X0,k1_zfmisc_1(k1_zfmisc_1(sK0))) | ~v1_xboole_0(X0) <- (31) [avatar component clause 567] {a:3,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
570. 32 <=> ! [X1] : ~r2_hidden(X1,sK1) [avatar definition]
571. ~r2_hidden(X1,sK1) <- (32) [avatar component clause 570] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
572. 31 | 32 | ~20 [avatar split clause 565,374,570,567]
700. ~v1_xboole_0(X0) <- (21, 31) [subsumption resolution 568,422] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
717. $false <- (21, 31) [resolution 700,145] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
720. ~21 | ~31 [avatar contradiction clause 717]
743. v1_xboole_0(sK1) <- (32) [resolution 571,227] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
749. $false <- (~18, 32) [subsumption resolution 743,348] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
750. 18 | ~32 [avatar contradiction clause 749]
751. $false [avatar sat refutation 163,168,173,183,370,377,395,572,720,750] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:303,thDist:-303}
% SZS output end Proof for setfam_1__t32_setfam_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 40
% Introduced skolems: 11
% 
% >>> Saturation
% Initial clauses: 49
% Generated clauses: 535
% Active clauses: 127
% Passive clauses: 307
% Final active clauses: 85
% Final passive clauses: 173
% Main loop iterations started: 134
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 40
% Fw demodulations: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 27
% Forward subsumptions: 111
% Backward subsumptions: 38
% 
% >>> Generating Inferences
% Binary resolution: 311
% Unit resulting resolution: 8
% Forward superposition: 23
% Backward superposition: 26
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% Sat splitting refutations: 13
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 46
% SAT solver unit clauses: 8
% SAT solver binary clauses: 18
% 
% Memory used [KB]: 5373
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% term sharing: 0.001 s
% forward subsumption resolution: 0.001 s
% forward demodulation: 0.001 s
% resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.006 s ( own 0.002 s ) 

