i: [1,26,0,8,3,1,0,0] t4_gate_1
a: [3,40,0,2,1,26]
s: 40
a: [3,35,0,4,1,26]
s: 35
i: [1,31,0,2,1,0,0,1] rc2_xboole_0
a: [3,51,0,2,1,31]
s: 51
i: [1,32,0,2,1,0,0,1] rc1_xboole_0
a: [3,56,0,2,1,32]
s: 56
i: [1,29,0,6,2,0,0,1] fc4_gate_1
s: 29
d: [2,58,1,0,0,2,52,40,35,29]
e: 58
d: [2,59,1,2,1,1,40,29,35]
d: [2,61,1,0,0,2,21,59,40]
e: 61
r: 40
a: [3,39,0,2,1,26]
s: 39
a: [3,44,0,2,1,26]
s: 44
i: [1,27,0,7,1,0,0,255] commutativity_k2_gate_1
s: 27
d: [2,64,1,6,2,0,44,29,27]
s: 64
d: [2,66,2,0,0,2,52,44,35,64]
e: 66
d: [2,67,2,2,1,1,40,64,35]
d: [2,71,2,0,0,2,21,67,44]
e: 71
r: 35
r: 44
a: [3,43,0,2,1,26]
s: 43
a: [3,36,0,4,1,26]
s: 36
d: [2,65,1,6,2,0,44,29,27]
s: 65
i: [1,28,0,7,1,0,0,255] commutativity_k2_gate_1
s: 28
i: [1,30,0,8,3,0,0,1] fc3_gate_1
s: 30
d: [2,73,1,0,0,3,52,39,43,36,30]
e: 73
d: [2,76,1,4,2,1,40,30,36]
d: [2,80,1,2,1,2,21,76,43]
d: [2,81,1,0,0,3,21,80,39]
e: 81
i: [1,25,0,6,2,1,0,0] t4_gate_1
i: [1,24,0,6,2,1,0,0] t4_gate_1
f: 81,71,59,24,25,26
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t4_gate_1
% SZS output start Proof for gate_1__t4_gate_1
1. ! [X0] : ! [X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) <=> (~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input t4_gate_1]
2. ~! [X0] : ! [X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) <=> (~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [negated conjecture 1]
3. ! [X0,X1] : k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [input commutativity_k2_gate_1]
7. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_gate_1(X0,X1))) [input fc3_gate_1]
8. ! [X0,X1] : (v1_xboole_0(X0) => v1_xboole_0(k2_gate_1(X0,X1))) [input fc4_gate_1]
11. ~! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) <=> (~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 2]
12. ? [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) <~> (~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [ennf transformation 11]
13. ! [X0,X1] : (v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 8]
14. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 7]
15. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 14]
16. ? [X0,X1] : (((v1_xboole_0(X1) | v1_xboole_0(X0)) | v1_xboole_0(k2_gate_1(X0,X1))) & ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k2_gate_1(X0,X1)))) [nnf transformation 12]
17. ? [X0,X1] : ((v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k2_gate_1(X0,X1))) & ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k2_gate_1(X0,X1)))) [flattening 16]
18. ? [X0,X1] : ((v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k2_gate_1(X0,X1))) & ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k2_gate_1(X0,X1)))) => ((v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k2_gate_1(sK0,sK1))) & ((~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k2_gate_1(sK0,sK1)))) [choice axiom]
19. (v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k2_gate_1(sK0,sK1))) & ((~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k2_gate_1(sK0,sK1))) [skolemisation 17,18]
24. ~v1_xboole_0(sK0) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) [cnf transformation 19] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
25. ~v1_xboole_0(sK1) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) [cnf transformation 19] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
26. v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k2_gate_1(sK0,sK1)) [cnf transformation 19] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
29. v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 13] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
30. ~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 15] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. 1 <=> v1_xboole_0(k2_gate_1(sK0,sK1)) [avatar definition]
35. ~v1_xboole_0(k2_gate_1(sK0,sK1)) <- (~1) [avatar component clause 34] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
36. v1_xboole_0(k2_gate_1(sK0,sK1)) <- (1) [avatar component clause 34] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
38. 2 <=> v1_xboole_0(sK0) [avatar definition]
39. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 38] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
42. 3 <=> v1_xboole_0(sK1) [avatar definition]
43. ~v1_xboole_0(sK1) <- (~3) [avatar component clause 42] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
44. v1_xboole_0(sK1) <- (3) [avatar component clause 42] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
45. 1 | 2 | 3 [avatar split clause 26,42,38,34]
46. ~1 | ~3 [avatar split clause 25,42,34]
47. ~1 | ~2 [avatar split clause 24,38,34]
59. ~v1_xboole_0(sK0) <- (~1) [resolution 29,35] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
63. ~2 | 1 [avatar split clause 59,34,38]
64. v1_xboole_0(k2_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 29,27] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
67. ~v1_xboole_0(sK1) <- (~1) [resolution 64,35] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
71. $false <- (~1, 3) [subsumption resolution 67,44] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
72. 1 | ~3 [avatar contradiction clause 71]
76. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1) [resolution 30,36] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
80. v1_xboole_0(sK0) <- (1, ~3) [subsumption resolution 76,43] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
81. $false <- (1, ~2, ~3) [subsumption resolution 80,39] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
82. ~1 | 2 | 3 [avatar contradiction clause 81]
83. $false [avatar sat refutation 45,46,47,63,72,82] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:52,thDist:-52}
% SZS output end Proof for gate_1__t4_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 9
% Generated clauses: 36
% Active clauses: 12
% Passive clauses: 14
% Final active clauses: 9
% Main loop iterations started: 14
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 5
% Unit resulting resolution: 3
% Forward superposition: 4
% Backward superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 2
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

