i: [1,39,0,6,1,1,0,0] t21_zfmisc_1
a: [3,56,0,6,1,39]
s: 56
i: [1,38,0,6,1,1,0,0] t21_zfmisc_1
a: [3,61,0,6,1,38]
s: 61
i: [1,40,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,66,0,2,1,40]
s: 66
i: [1,41,0,2,1,0,0,255] fc1_xboole_0
a: [3,71,0,2,1,41]
s: 71
i: [1,51,0,2,1,0,0,2] rc2_xboole_0
a: [3,81,0,2,1,51]
s: 81
i: [1,52,0,2,1,0,0,2] rc1_xboole_0
a: [3,86,0,2,1,52]
s: 86
i: [1,43,0,3,1,0,0,1] fc2_xboole_0
s: 43
i: [1,48,0,3,1,0,0,1] reflexivity_r1_tarski
s: 48
i: [1,42,0,3,1,0,0,255] d2_xboole_0
a: [3,76,0,3,1,42]
s: 76
r: 66
i: [1,49,0,4,1,0,0,1] fc3_xboole_0
s: 49
i: [1,45,0,5,2,0,0,255] t6_boole
s: 45
i: [1,44,0,6,1,0,0,1] t29_enumset1
s: 44
i: [1,46,0,7,3,0,0,1] t8_boole
s: 46
i: [1,50,0,7,1,0,0,1] commutativity_k2_tarski
s: 50
i: [1,47,0,9,2,0,0,1] t20_zfmisc_1
s: 47
d: [2,94,1,3,1,1,40,47,61]
a: [3,101,1,3,1,94]
s: 101
r: 61
r: 56
d: [2,103,0,6,1,2,23,61,101]
a: [3,108,0,6,1,103]
s: 108
d: [2,104,0,6,1,2,23,56,101]
a: [3,113,0,6,1,104]
s: 113
d: [2,89,1,7,3,0,44,45,45]
s: 89
d: [2,95,1,8,2,0,44,47,44]
s: 95
d: [2,96,1,9,2,0,44,47,50]
s: 96
d: [2,117,2,3,1,1,40,96,108]
a: [3,124,2,3,1,117]
s: 124
r: 113
r: 108
d: [2,126,0,6,1,2,23,113,124]
d: [2,128,0,0,0,2,21,126,44]
e: 128
f: 128,117,104,103,94,38,39
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t21_zfmisc_1
% SZS output start Proof for zfmisc_1__t21_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),k1_tarski(X2)) => k2_tarski(X0,X1) = k1_tarski(X2)) [input t21_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),k1_tarski(X2)) => k2_tarski(X0,X1) = k1_tarski(X2)) [negated conjecture 1]
4. ! [X0,X1] : k2_tarski(X0,X1) = k2_tarski(X1,X0) [input commutativity_k2_tarski]
16. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),k1_tarski(X2)) => X0 = X2) [input t20_zfmisc_1]
17. ! [X0] : k1_tarski(X0) = k2_tarski(X0,X0) [input t29_enumset1]
21. ~! [X0,X1,X2] : (r1_tarski(k2_tarski(X0,X1),k1_tarski(X2)) => k2_tarski(X0,X1) = k1_tarski(X2)) [flattening 2]
24. ! [X0,X1,X2] : (r1_tarski(k2_tarski(X0,X1),k1_tarski(X2)) => X0 = X2) [flattening 16]
28. ? [X0,X1,X2] : (k2_tarski(X0,X1) != k1_tarski(X2) & r1_tarski(k2_tarski(X0,X1),k1_tarski(X2))) [ennf transformation 21]
31. ! [X0,X1,X2] : (X0 = X2 | ~r1_tarski(k2_tarski(X0,X1),k1_tarski(X2))) [ennf transformation 24]
32. ? [X0,X1,X2] : (k2_tarski(X0,X1) != k1_tarski(X2) & r1_tarski(k2_tarski(X0,X1),k1_tarski(X2))) => (k2_tarski(sK0,sK1) != k1_tarski(sK2) & r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2))) [choice axiom]
33. k2_tarski(sK0,sK1) != k1_tarski(sK2) & r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2)) [skolemisation 28,32]
38. r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2)) [cnf transformation 33] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. k2_tarski(sK0,sK1) != k1_tarski(sK2) [cnf transformation 33] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
44. k1_tarski(X0) = k2_tarski(X0,X0) [cnf transformation 17] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. ~r1_tarski(k2_tarski(X0,X1),k1_tarski(X2)) | X0 = X2 [cnf transformation 31] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. k2_tarski(X0,X1) = k2_tarski(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. 1 <=> k2_tarski(sK0,sK1) = k1_tarski(sK2) [avatar definition]
56. k2_tarski(sK0,sK1) != k1_tarski(sK2) <- (~1) [avatar component clause 54] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
57. ~1 [avatar split clause 39,54]
59. 2 <=> r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2)) [avatar definition]
61. r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK2)) <- (2) [avatar component clause 59] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. 2 [avatar split clause 38,59]
94. sK0 = sK2 <- (2) [resolution 47,61] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
96. ~r1_tarski(k2_tarski(X3,X2),k1_tarski(X4)) | X2 = X4 [superposition 47,50] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
99. 8 <=> sK0 = sK2 [avatar definition]
101. sK0 = sK2 <- (8) [avatar component clause 99] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
102. 8 | ~2 [avatar split clause 94,59,99]
103. r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK0)) <- (2, 8) [backward demodulation 61,101] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
104. k2_tarski(sK0,sK1) != k1_tarski(sK0) <- (~1, 8) [backward demodulation 56,101] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
106. 9 <=> r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK0)) [avatar definition]
108. r1_tarski(k2_tarski(sK0,sK1),k1_tarski(sK0)) <- (9) [avatar component clause 106] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
109. 9 | ~2 | ~8 [avatar split clause 103,99,59,106]
111. 10 <=> k2_tarski(sK0,sK1) = k1_tarski(sK0) [avatar definition]
113. k2_tarski(sK0,sK1) != k1_tarski(sK0) <- (~10) [avatar component clause 111] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
114. ~10 | 1 | ~8 [avatar split clause 104,99,54,111]
117. sK0 = sK1 <- (9) [resolution 96,108] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
122. 11 <=> sK0 = sK1 [avatar definition]
124. sK0 = sK1 <- (11) [avatar component clause 122] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
125. 11 | ~9 [avatar split clause 117,106,122]
126. k1_tarski(sK0) != k2_tarski(sK0,sK0) <- (~10, 11) [backward demodulation 113,124] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
128. $false <- (~10, 11) [subsumption resolution 126,44] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
129. 10 | ~11 [avatar contradiction clause 128]
131. $false [avatar sat refutation 57,62,102,109,114,125,129] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:41,thDist:-41}
% SZS output end Proof for zfmisc_1__t21_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 19
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 48
% Active clauses: 21
% Passive clauses: 23
% Final active clauses: 16
% Final passive clauses: 1
% Main loop iterations started: 22
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 1
% Fw demodulations: 1
% Bw demodulations: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Equational tautologies: 2
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 4
% Forward superposition: 6
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 7
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

