i: [1,31,0,2,1,1,0,0] t44_wellord1
a: [3,56,0,2,1,31]
s: 56
i: [1,29,0,2,1,1,0,0] t44_wellord1
a: [3,66,0,2,1,29]
s: 66
i: [1,28,0,2,1,1,0,0] t44_wellord1
a: [3,71,0,2,1,28]
s: 71
i: [1,27,0,2,1,1,0,0] t44_wellord1
a: [3,76,0,2,1,27]
s: 76
i: [1,26,0,2,1,1,0,0] t44_wellord1
a: [3,81,0,2,1,26]
s: 81
i: [1,25,0,2,1,1,0,0] t44_wellord1
a: [3,86,0,2,1,25]
s: 86
i: [1,30,0,4,1,1,0,0] t44_wellord1
a: [3,61,0,4,1,30]
s: 61
i: [1,52,0,2,1,0,0,1] rc1_funct_1
a: [3,91,0,2,1,52]
s: 91
i: [1,51,0,2,1,0,0,1] rc1_funct_1
a: [3,96,0,2,1,51]
s: 96
i: [1,32,0,6,3,0,0,1] d4_wellord1
s: 32
i: [1,33,0,6,3,0,0,1] d4_wellord1
s: 33
d: [2,98,1,4,2,1,40,32,66]
d: [2,99,1,2,1,2,21,98,86]
a: [3,103,1,2,1,99]
s: 103
i: [1,34,0,6,3,0,0,1] d4_wellord1
s: 34
d: [2,105,1,4,2,1,40,33,66]
d: [2,106,1,2,1,2,21,105,86]
a: [3,110,1,2,1,106]
s: 110
i: [1,35,0,6,3,0,0,1] d4_wellord1
s: 35
i: [1,36,0,6,3,0,0,1] d4_wellord1
s: 36
d: [2,112,1,4,2,1,40,34,66]
d: [2,113,1,2,1,2,21,112,86]
a: [3,117,1,2,1,113]
s: 117
i: [1,46,0,6,3,0,0,1] cc1_wellord1
s: 46
d: [2,119,1,4,2,1,40,35,66]
d: [2,120,1,2,1,2,21,119,86]
a: [3,124,1,2,1,120]
s: 124
i: [1,47,0,6,3,0,0,1] cc1_wellord1
s: 47
i: [1,48,0,6,3,0,0,1] cc1_wellord1
s: 48
d: [2,126,1,4,2,1,40,36,66]
d: [2,127,1,2,1,2,21,126,86]
a: [3,131,1,2,1,127]
s: 131
i: [1,49,0,6,3,0,0,1] cc1_wellord1
s: 49
i: [1,50,0,6,3,0,0,1] cc1_wellord1
s: 50
i: [1,37,0,14,7,0,0,1] d4_wellord1
s: 37
i: [1,44,0,14,7,0,0,1] cc2_wellord1
s: 44
i: [1,38,0,16,7,0,0,1] t43_wellord1
s: 38
i: [1,39,0,16,7,0,0,1] t43_wellord1
s: 39
d: [2,134,1,12,6,1,40,38,61]
d: [2,135,1,10,5,2,21,134,103]
d: [2,136,1,8,4,3,21,135,71]
d: [2,137,1,6,3,4,21,136,76]
d: [2,138,1,4,2,5,21,137,81]
d: [2,139,1,2,1,6,21,138,86]
a: [3,143,1,2,1,139]
s: 143
i: [1,40,0,16,7,0,0,1] t43_wellord1
s: 40
i: [1,41,0,16,7,0,0,1] t43_wellord1
s: 41
d: [2,145,1,12,6,1,40,39,61]
d: [2,146,1,10,5,2,21,145,110]
d: [2,147,1,8,4,3,21,146,71]
d: [2,148,1,6,3,4,21,147,76]
d: [2,149,1,4,2,5,21,148,81]
d: [2,150,1,2,1,6,21,149,86]
a: [3,154,1,2,1,150]
s: 154
i: [1,42,0,16,7,0,0,1] t43_wellord1
s: 42
d: [2,156,1,12,6,1,40,40,61]
d: [2,157,1,10,5,2,21,156,124]
d: [2,158,1,8,4,3,21,157,71]
d: [2,159,1,6,3,4,21,158,76]
d: [2,160,1,4,2,5,21,159,81]
d: [2,161,1,2,1,6,21,160,86]
a: [3,165,1,2,1,161]
s: 165
d: [2,167,1,12,6,1,40,41,61]
d: [2,168,1,10,5,2,21,167,117]
d: [2,169,1,8,4,3,21,168,71]
d: [2,170,1,6,3,4,21,169,76]
d: [2,171,1,4,2,5,21,170,81]
d: [2,172,1,2,1,6,21,171,86]
a: [3,176,1,2,1,172]
s: 176
d: [2,178,1,12,6,1,40,42,61]
d: [2,179,1,10,5,2,21,178,131]
d: [2,180,1,8,4,3,21,179,71]
d: [2,181,1,6,3,4,21,180,76]
d: [2,182,1,4,2,5,21,181,81]
d: [2,183,1,2,1,6,21,182,86]
a: [3,187,1,2,1,183]
s: 187
d: [2,189,2,0,0,7,52,81,143,154,176,165,56,187,37]
e: 189
d: [2,190,2,12,6,1,40,187,37]
d: [2,192,2,10,5,2,21,190,56]
d: [2,193,2,8,4,3,21,192,165]
d: [2,194,2,6,3,4,21,193,176]
d: [2,195,2,4,2,5,21,194,154]
d: [2,196,2,2,1,6,21,195,143]
d: [2,197,2,0,0,7,21,196,81]
e: 197
f: 197,183,172,161,150,139,127,120,113,106,99,25,26,27,28,29,30,31
% Refutation found. Thanks to Tanya!
% SZS status Theorem for wellord1__t44_wellord1
% SZS output start Proof for wellord1__t44_wellord1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r3_wellord1(X0,X1,X2) & v2_wellord1(X0)) => v2_wellord1(X1))))) [input t44_wellord1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r3_wellord1(X0,X1,X2) & v2_wellord1(X0)) => v2_wellord1(X1))))) [negated conjecture 1]
5. ! [X0] : (v1_relat_1(X0) => (v2_wellord1(X0) <=> (v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d4_wellord1]
7. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => (r3_wellord1(X0,X1,X2) => ((v1_wellord1(X0) => v1_wellord1(X1)) & (v4_relat_2(X0) => v4_relat_2(X1)) & (v6_relat_2(X0) => v6_relat_2(X1)) & (v8_relat_2(X0) => v8_relat_2(X1)) & (v1_relat_2(X0) => v1_relat_2(X1))))))) [input t43_wellord1]
8. ? [X0] : (? [X1] : (? [X2] : ((~v2_wellord1(X1) & (r3_wellord1(X0,X1,X2) & v2_wellord1(X0))) & (v1_funct_1(X2) & v1_relat_1(X2))) & v1_relat_1(X1)) & v1_relat_1(X0)) [ennf transformation 2]
9. ? [X0] : (? [X1] : (? [X2] : (~v2_wellord1(X1) & r3_wellord1(X0,X1,X2) & v2_wellord1(X0) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_relat_1(X1)) & v1_relat_1(X0)) [flattening 8]
10. ! [X0] : ((v2_wellord1(X0) <=> (v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 5]
11. ! [X0] : (! [X1] : (! [X2] : ((((v1_wellord1(X1) | ~v1_wellord1(X0)) & (v4_relat_2(X1) | ~v4_relat_2(X0)) & (v6_relat_2(X1) | ~v6_relat_2(X0)) & (v8_relat_2(X1) | ~v8_relat_2(X0)) & (v1_relat_2(X1) | ~v1_relat_2(X0))) | ~r3_wellord1(X0,X1,X2)) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 7]
12. ! [X0] : (! [X1] : (! [X2] : (((v1_wellord1(X1) | ~v1_wellord1(X0)) & (v4_relat_2(X1) | ~v4_relat_2(X0)) & (v6_relat_2(X1) | ~v6_relat_2(X0)) & (v8_relat_2(X1) | ~v8_relat_2(X0)) & (v1_relat_2(X1) | ~v1_relat_2(X0))) | ~r3_wellord1(X0,X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 11]
17. ? [X0] : (? [X1] : (? [X2] : (~v2_wellord1(X1) & r3_wellord1(X0,X1,X2) & v2_wellord1(X0) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_relat_1(X1)) & v1_relat_1(X0)) => (? [X1] : (? [X2] : (~v2_wellord1(X1) & r3_wellord1(sK0,X1,X2) & v2_wellord1(sK0) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_relat_1(X1)) & v1_relat_1(sK0)) [choice axiom]
18. ? [X1] : (? [X2] : (~v2_wellord1(X1) & r3_wellord1(sK0,X1,X2) & v2_wellord1(sK0) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_relat_1(X1)) => (? [X2] : (~v2_wellord1(sK1) & r3_wellord1(sK0,sK1,X2) & v2_wellord1(sK0) & v1_funct_1(X2) & v1_relat_1(X2)) & v1_relat_1(sK1)) [choice axiom]
19. ? [X2] : (~v2_wellord1(sK1) & r3_wellord1(sK0,sK1,X2) & v2_wellord1(sK0) & v1_funct_1(X2) & v1_relat_1(X2)) => (~v2_wellord1(sK1) & r3_wellord1(sK0,sK1,sK2) & v2_wellord1(sK0) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
20. ((~v2_wellord1(sK1) & r3_wellord1(sK0,sK1,sK2) & v2_wellord1(sK0) & v1_funct_1(sK2) & v1_relat_1(sK2)) & v1_relat_1(sK1)) & v1_relat_1(sK0) [skolemisation 9,19,18,17]
21. ! [X0] : (((v2_wellord1(X0) | (~v1_wellord1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_wellord1(X0))) | ~v1_relat_1(X0)) [nnf transformation 10]
22. ! [X0] : (((v2_wellord1(X0) | ~v1_wellord1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v1_wellord1(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_wellord1(X0))) | ~v1_relat_1(X0)) [flattening 21]
25. v1_relat_1(sK0) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
26. v1_relat_1(sK1) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
27. v1_relat_1(sK2) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
28. v1_funct_1(sK2) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
29. v2_wellord1(sK0) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
30. r3_wellord1(sK0,sK1,sK2) [cnf transformation 20] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
31. ~v2_wellord1(sK1) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
32. ~v2_wellord1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. ~v2_wellord1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. ~v2_wellord1(X0) | v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
35. ~v2_wellord1(X0) | v6_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
36. ~v2_wellord1(X0) | v1_wellord1(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. ~v1_wellord1(X0) | v2_wellord1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
38. ~r3_wellord1(X0,X1,X2) | ~v1_relat_2(X0) | v1_relat_2(X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 12] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
39. ~r3_wellord1(X0,X1,X2) | ~v8_relat_2(X0) | v8_relat_2(X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 12] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
40. ~r3_wellord1(X0,X1,X2) | ~v6_relat_2(X0) | v6_relat_2(X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 12] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
41. ~r3_wellord1(X0,X1,X2) | ~v4_relat_2(X0) | v4_relat_2(X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 12] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. ~r3_wellord1(X0,X1,X2) | ~v1_wellord1(X0) | v1_wellord1(X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 12] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. 1 <=> v2_wellord1(sK1) [avatar definition]
56. ~v2_wellord1(sK1) <- (~1) [avatar component clause 54] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
57. ~1 [avatar split clause 31,54]
59. 2 <=> r3_wellord1(sK0,sK1,sK2) [avatar definition]
61. r3_wellord1(sK0,sK1,sK2) <- (2) [avatar component clause 59] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
62. 2 [avatar split clause 30,59]
64. 3 <=> v2_wellord1(sK0) [avatar definition]
66. v2_wellord1(sK0) <- (3) [avatar component clause 64] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
67. 3 [avatar split clause 29,64]
69. 4 <=> v1_funct_1(sK2) [avatar definition]
71. v1_funct_1(sK2) <- (4) [avatar component clause 69] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
72. 4 [avatar split clause 28,69]
74. 5 <=> v1_relat_1(sK2) [avatar definition]
76. v1_relat_1(sK2) <- (5) [avatar component clause 74] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
77. 5 [avatar split clause 27,74]
79. 6 <=> v1_relat_1(sK1) [avatar definition]
81. v1_relat_1(sK1) <- (6) [avatar component clause 79] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
82. 6 [avatar split clause 26,79]
84. 7 <=> v1_relat_1(sK0) [avatar definition]
86. v1_relat_1(sK0) <- (7) [avatar component clause 84] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
87. 7 [avatar split clause 25,84]
98. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (3) [resolution 32,66] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
99. v1_relat_2(sK0) <- (3, 7) [subsumption resolution 98,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
101. 10 <=> v1_relat_2(sK0) [avatar definition]
103. v1_relat_2(sK0) <- (10) [avatar component clause 101] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
104. 10 | ~3 | ~7 [avatar split clause 99,84,64,101]
105. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (3) [resolution 33,66] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
106. v8_relat_2(sK0) <- (3, 7) [subsumption resolution 105,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
108. 11 <=> v8_relat_2(sK0) [avatar definition]
110. v8_relat_2(sK0) <- (11) [avatar component clause 108] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
111. 11 | ~3 | ~7 [avatar split clause 106,84,64,108]
112. v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (3) [resolution 34,66] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
113. v4_relat_2(sK0) <- (3, 7) [subsumption resolution 112,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
115. 12 <=> v4_relat_2(sK0) [avatar definition]
117. v4_relat_2(sK0) <- (12) [avatar component clause 115] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
118. 12 | ~3 | ~7 [avatar split clause 113,84,64,115]
119. v6_relat_2(sK0) | ~v1_relat_1(sK0) <- (3) [resolution 35,66] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
120. v6_relat_2(sK0) <- (3, 7) [subsumption resolution 119,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
122. 13 <=> v6_relat_2(sK0) [avatar definition]
124. v6_relat_2(sK0) <- (13) [avatar component clause 122] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
125. 13 | ~3 | ~7 [avatar split clause 120,84,64,122]
126. v1_wellord1(sK0) | ~v1_relat_1(sK0) <- (3) [resolution 36,66] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
127. v1_wellord1(sK0) <- (3, 7) [subsumption resolution 126,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
129. 14 <=> v1_wellord1(sK0) [avatar definition]
131. v1_wellord1(sK0) <- (14) [avatar component clause 129] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
132. 14 | ~3 | ~7 [avatar split clause 127,84,64,129]
134. ~v1_relat_2(sK0) | v1_relat_2(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2) [resolution 38,61] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
135. v1_relat_2(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 10) [subsumption resolution 134,103] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
136. v1_relat_2(sK1) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 10) [subsumption resolution 135,71] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
137. v1_relat_2(sK1) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 10) [subsumption resolution 136,76] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
138. v1_relat_2(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 6, 10) [subsumption resolution 137,81] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
139. v1_relat_2(sK1) <- (2, 4, 5, 6, 7, 10) [subsumption resolution 138,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
141. 15 <=> v1_relat_2(sK1) [avatar definition]
143. v1_relat_2(sK1) <- (15) [avatar component clause 141] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
144. 15 | ~2 | ~4 | ~5 | ~6 | ~7 | ~10 [avatar split clause 139,101,84,79,74,69,59,141]
145. ~v8_relat_2(sK0) | v8_relat_2(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2) [resolution 39,61] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
146. v8_relat_2(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 11) [subsumption resolution 145,110] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
147. v8_relat_2(sK1) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 11) [subsumption resolution 146,71] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
148. v8_relat_2(sK1) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 11) [subsumption resolution 147,76] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
149. v8_relat_2(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 6, 11) [subsumption resolution 148,81] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
150. v8_relat_2(sK1) <- (2, 4, 5, 6, 7, 11) [subsumption resolution 149,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
152. 16 <=> v8_relat_2(sK1) [avatar definition]
154. v8_relat_2(sK1) <- (16) [avatar component clause 152] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
155. 16 | ~2 | ~4 | ~5 | ~6 | ~7 | ~11 [avatar split clause 150,108,84,79,74,69,59,152]
156. ~v6_relat_2(sK0) | v6_relat_2(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2) [resolution 40,61] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
157. v6_relat_2(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 13) [subsumption resolution 156,124] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
158. v6_relat_2(sK1) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 13) [subsumption resolution 157,71] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
159. v6_relat_2(sK1) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 13) [subsumption resolution 158,76] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
160. v6_relat_2(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 6, 13) [subsumption resolution 159,81] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
161. v6_relat_2(sK1) <- (2, 4, 5, 6, 7, 13) [subsumption resolution 160,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
163. 17 <=> v6_relat_2(sK1) [avatar definition]
165. v6_relat_2(sK1) <- (17) [avatar component clause 163] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
166. 17 | ~2 | ~4 | ~5 | ~6 | ~7 | ~13 [avatar split clause 161,122,84,79,74,69,59,163]
167. ~v4_relat_2(sK0) | v4_relat_2(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2) [resolution 41,61] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
168. v4_relat_2(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 12) [subsumption resolution 167,117] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
169. v4_relat_2(sK1) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 12) [subsumption resolution 168,71] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
170. v4_relat_2(sK1) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 12) [subsumption resolution 169,76] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
171. v4_relat_2(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 6, 12) [subsumption resolution 170,81] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
172. v4_relat_2(sK1) <- (2, 4, 5, 6, 7, 12) [subsumption resolution 171,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
174. 18 <=> v4_relat_2(sK1) [avatar definition]
176. v4_relat_2(sK1) <- (18) [avatar component clause 174] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
177. 18 | ~2 | ~4 | ~5 | ~6 | ~7 | ~12 [avatar split clause 172,115,84,79,74,69,59,174]
178. ~v1_wellord1(sK0) | v1_wellord1(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2) [resolution 42,61] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
179. v1_wellord1(sK1) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 14) [subsumption resolution 178,131] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
180. v1_wellord1(sK1) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 14) [subsumption resolution 179,71] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
181. v1_wellord1(sK1) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 14) [subsumption resolution 180,76] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
182. v1_wellord1(sK1) | ~v1_relat_1(sK0) <- (2, 4, 5, 6, 14) [subsumption resolution 181,81] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
183. v1_wellord1(sK1) <- (2, 4, 5, 6, 7, 14) [subsumption resolution 182,86] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
185. 19 <=> v1_wellord1(sK1) [avatar definition]
187. v1_wellord1(sK1) <- (19) [avatar component clause 185] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
188. 19 | ~2 | ~4 | ~5 | ~6 | ~7 | ~14 [avatar split clause 183,129,84,79,74,69,59,185]
190. v2_wellord1(sK1) | ~v6_relat_2(sK1) | ~v4_relat_2(sK1) | ~v8_relat_2(sK1) | ~v1_relat_2(sK1) | ~v1_relat_1(sK1) <- (19) [resolution 187,37] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
192. ~v6_relat_2(sK1) | ~v4_relat_2(sK1) | ~v8_relat_2(sK1) | ~v1_relat_2(sK1) | ~v1_relat_1(sK1) <- (~1, 19) [subsumption resolution 190,56] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
193. ~v4_relat_2(sK1) | ~v8_relat_2(sK1) | ~v1_relat_2(sK1) | ~v1_relat_1(sK1) <- (~1, 17, 19) [subsumption resolution 192,165] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
194. ~v8_relat_2(sK1) | ~v1_relat_2(sK1) | ~v1_relat_1(sK1) <- (~1, 17, 18, 19) [subsumption resolution 193,176] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
195. ~v1_relat_2(sK1) | ~v1_relat_1(sK1) <- (~1, 16, 17, 18, 19) [subsumption resolution 194,154] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
196. ~v1_relat_1(sK1) <- (~1, 15, 16, 17, 18, 19) [subsumption resolution 195,143] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
197. $false <- (~1, 6, 15, 16, 17, 18, 19) [subsumption resolution 196,81] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
198. 1 | ~6 | ~15 | ~16 | ~17 | ~18 | ~19 [avatar contradiction clause 197]
199. $false [avatar sat refutation 57,62,67,72,77,82,87,104,111,118,125,132,144,155,166,177,188,198] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:181,thDist:-181}
% SZS output end Proof for wellord1__t44_wellord1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 6
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 28
% Generated clauses: 96
% Active clauses: 30
% Passive clauses: 36
% Final active clauses: 30
% Main loop iterations started: 36
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 36
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 7
% 
% >>> Generating Inferences
% Binary resolution: 12
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 9
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

