i: [1,52,0,2,1,1,0,0] t67_xxreal_2
a: [3,82,0,2,1,52]
s: 82
i: [1,51,0,2,1,1,0,0] t67_xxreal_2
a: [3,87,0,2,1,51]
s: 87
i: [1,71,0,2,1,0,0,2] rc1_xxreal_0
a: [3,103,0,2,1,71]
s: 103
i: [1,53,0,10,1,1,0,0] t67_xxreal_2
a: [3,77,0,10,1,53]
s: 77
i: [1,54,0,5,2,0,0,1] dt_k1_xxreal_2
s: 54
i: [1,62,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 62
i: [1,66,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,94,0,5,2,66]
s: 94
i: [1,61,0,6,2,0,0,2] existence_m1_xxreal_2
s: 61
i: [1,64,0,6,2,0,0,1] fc33_membered
s: 64
i: [1,65,0,6,2,0,0,1] fc34_membered
s: 65
i: [1,60,0,7,3,0,0,2] dt_m1_xxreal_2
s: 60
d: [2,107,1,7,3,0,40,60,61]
d: [2,108,1,5,2,0,18,107]
s: 108
i: [1,63,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 63
i: [1,68,0,9,3,0,0,1] idempotence_k3_xxreal_0
a: [3,98,0,7,2,68]
s: 98
i: [1,67,0,8,3,0,0,1] fc5_xxreal_0
s: 67
i: [1,73,0,9,3,0,0,1] d3_xxreal_2
s: 73
d: [2,115,0,6,2,0,21,73,54]
s: 115
i: [1,70,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 70
d: [2,118,1,21,3,1,40,70,77]
a: [3,124,1,6,1,118]
s: 124
i: [1,69,0,11,3,0,0,1] commutativity_k3_xxreal_0
s: 69
i: [1,72,0,12,4,0,0,1] d3_xxreal_2
s: 72
d: [2,134,2,6,2,1,40,124,67]
a: [3,138,2,3,1,134]
s: 138
d: [2,149,3,0,0,2,52,87,138,54]
e: 149
d: [2,150,3,2,1,1,40,138,54]
d: [2,152,3,0,0,2,21,150,87]
e: 152
r: 138
d: [2,148,0,9,3,0,21,72,54]
s: 148
a: [3,137,2,3,1,134]
s: 137
i: [1,58,0,16,5,0,0,1] d3_xxreal_2
s: 58
d: [2,164,0,14,4,0,21,58,60]
s: 164
a: [3,142,2,3,1,134]
s: 142
d: [2,167,3,0,0,2,52,82,142,54]
e: 167
d: [2,168,3,2,1,1,40,142,54]
d: [2,170,3,0,0,2,21,168,82]
e: 170
r: 124
r: 142
i: [1,59,0,16,5,0,0,1] d3_xxreal_2
s: 59
a: [3,141,2,3,1,134]
s: 141
d: [2,172,0,14,4,0,21,59,60]
s: 172
i: [1,55,0,17,5,0,0,2] t65_xxreal_2
s: 55
d: [2,154,1,13,2,1,40,148,77]
a: [3,158,1,4,1,154]
s: 158
d: [2,190,2,0,0,2,52,82,158,65]
e: 190
d: [2,191,2,0,0,2,52,87,158,64]
e: 191
d: [2,192,2,2,1,1,40,158,65]
d: [2,196,2,0,0,2,21,192,82]
e: 196
d: [2,193,2,2,1,1,40,158,64]
d: [2,198,2,0,0,2,21,193,87]
e: 198
r: 158
a: [3,157,1,4,1,154]
s: 157
a: [3,128,1,5,1,118]
s: 128
d: [2,200,2,0,0,2,52,157,128,54]
e: 200
d: [2,201,2,4,1,1,40,128,54]
d: [2,203,2,0,0,2,21,201,157]
e: 203
r: 128
a: [3,127,1,5,1,118]
s: 127
a: [3,123,1,6,1,118]
s: 123
a: [3,162,1,9,1,154]
s: 162
a: [3,132,1,10,1,118]
s: 132
d: [2,205,2,12,4,1,40,162,55]
d: [2,206,2,10,3,2,21,205,82]
d: [2,207,2,8,2,3,21,206,87]
a: [3,211,2,4,1,207]
s: 211
d: [2,217,3,0,0,2,52,87,211,115]
e: 217
d: [2,218,3,2,1,1,40,211,115]
d: [2,220,3,0,0,2,21,218,87]
e: 220
r: 211
d: [2,165,1,15,5,0,40,164,60]
d: [2,166,1,13,4,0,18,165]
s: 166
d: [2,180,1,15,5,0,44,55,62]
d: [2,183,1,13,4,0,18,180]
s: 183
a: [3,210,2,4,1,207]
s: 210
d: [2,176,1,18,6,0,40,55,60]
d: [2,184,1,14,5,0,21,176,65]
s: 184
a: [3,215,2,4,1,207]
s: 215
d: [2,243,3,0,0,2,52,82,215,115]
e: 243
d: [2,244,3,2,1,1,40,215,115]
d: [2,246,3,0,0,2,21,244,82]
e: 246
f: 246,220,207,198,154,51,52,53
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_2__t67_xxreal_2
% SZS output start Proof for xxreal_2__t67_xxreal_2
1. ! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => r1_xxreal_0(k1_xxreal_2(k3_xboole_0(X0,X1)),k3_xxreal_0(k1_xxreal_2(X0),k1_xxreal_2(X1))))) [input t67_xxreal_2]
2. ~! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => r1_xxreal_0(k1_xxreal_2(k3_xboole_0(X0,X1)),k3_xxreal_0(k1_xxreal_2(X0),k1_xxreal_2(X1))))) [negated conjecture 1]
6. ! [X0] : (v2_membered(X0) => ! [X1] : (v1_xxreal_0(X1) => (k1_xxreal_2(X0) = X1 <=> (! [X2] : (m1_xxreal_2(X2,X0) => r1_xxreal_0(X1,X2)) & m1_xxreal_2(X1,X0))))) [input d3_xxreal_2]
7. ! [X0] : (v2_membered(X0) => v1_xxreal_0(k1_xxreal_2(X0))) [input dt_k1_xxreal_2]
12. ! [X0,X1] : (v2_membered(X0) => v2_membered(k3_xboole_0(X0,X1))) [input fc33_membered]
19. ! [X0] : (v2_membered(X0) => ! [X1] : (v2_membered(X1) => ! [X2] : (m1_xxreal_2(X2,X0) => ! [X3] : (m1_xxreal_2(X3,X1) => m1_xxreal_2(k3_xxreal_0(X2,X3),k3_xboole_0(X0,X1)))))) [input t65_xxreal_2]
21. ? [X0] : (? [X1] : (~r1_xxreal_0(k1_xxreal_2(k3_xboole_0(X0,X1)),k3_xxreal_0(k1_xxreal_2(X0),k1_xxreal_2(X1))) & v2_membered(X1)) & v2_membered(X0)) [ennf transformation 2]
22. ! [X0] : (v1_xxreal_0(k1_xxreal_2(X0)) | ~v2_membered(X0)) [ennf transformation 7]
23. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (m1_xxreal_2(k3_xxreal_0(X2,X3),k3_xboole_0(X0,X1)) | ~m1_xxreal_2(X3,X1)) | ~m1_xxreal_2(X2,X0)) | ~v2_membered(X1)) | ~v2_membered(X0)) [ennf transformation 19]
24. ! [X0] : (! [X1] : ((k1_xxreal_2(X0) = X1 <=> (! [X2] : (r1_xxreal_0(X1,X2) | ~m1_xxreal_2(X2,X0)) & m1_xxreal_2(X1,X0))) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [ennf transformation 6]
27. ! [X0,X1] : (v2_membered(k3_xboole_0(X0,X1)) | ~v2_membered(X0)) [ennf transformation 12]
39. ? [X0] : (? [X1] : (~r1_xxreal_0(k1_xxreal_2(k3_xboole_0(X0,X1)),k3_xxreal_0(k1_xxreal_2(X0),k1_xxreal_2(X1))) & v2_membered(X1)) & v2_membered(X0)) => (? [X1] : (~r1_xxreal_0(k1_xxreal_2(k3_xboole_0(sK0,X1)),k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(X1))) & v2_membered(X1)) & v2_membered(sK0)) [choice axiom]
40. ? [X1] : (~r1_xxreal_0(k1_xxreal_2(k3_xboole_0(sK0,X1)),k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(X1))) & v2_membered(X1)) => (~r1_xxreal_0(k1_xxreal_2(k3_xboole_0(sK0,sK1)),k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1))) & v2_membered(sK1)) [choice axiom]
41. (~r1_xxreal_0(k1_xxreal_2(k3_xboole_0(sK0,sK1)),k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1))) & v2_membered(sK1)) & v2_membered(sK0) [skolemisation 21,40,39]
42. ! [X0] : (! [X1] : (((k1_xxreal_2(X0) = X1 | (? [X2] : (~r1_xxreal_0(X1,X2) & m1_xxreal_2(X2,X0)) | ~m1_xxreal_2(X1,X0))) & ((! [X2] : (r1_xxreal_0(X1,X2) | ~m1_xxreal_2(X2,X0)) & m1_xxreal_2(X1,X0)) | k1_xxreal_2(X0) != X1)) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [nnf transformation 24]
43. ! [X0] : (! [X1] : (((k1_xxreal_2(X0) = X1 | ? [X2] : (~r1_xxreal_0(X1,X2) & m1_xxreal_2(X2,X0)) | ~m1_xxreal_2(X1,X0)) & ((! [X2] : (r1_xxreal_0(X1,X2) | ~m1_xxreal_2(X2,X0)) & m1_xxreal_2(X1,X0)) | k1_xxreal_2(X0) != X1)) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [flattening 42]
44. ! [X0] : (! [X1] : (((k1_xxreal_2(X0) = X1 | ? [X2] : (~r1_xxreal_0(X1,X2) & m1_xxreal_2(X2,X0)) | ~m1_xxreal_2(X1,X0)) & ((! [X3] : (r1_xxreal_0(X1,X3) | ~m1_xxreal_2(X3,X0)) & m1_xxreal_2(X1,X0)) | k1_xxreal_2(X0) != X1)) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [rectify 43]
45. ! [X1,X0] : (? [X2] : (~r1_xxreal_0(X1,X2) & m1_xxreal_2(X2,X0)) => (~r1_xxreal_0(X1,sK2(X0,X1)) & m1_xxreal_2(sK2(X0,X1),X0))) [choice axiom]
46. ! [X0] : (! [X1] : (((k1_xxreal_2(X0) = X1 | (~r1_xxreal_0(X1,sK2(X0,X1)) & m1_xxreal_2(sK2(X0,X1),X0)) | ~m1_xxreal_2(X1,X0)) & ((! [X3] : (r1_xxreal_0(X1,X3) | ~m1_xxreal_2(X3,X0)) & m1_xxreal_2(X1,X0)) | k1_xxreal_2(X0) != X1)) | ~v1_xxreal_0(X1)) | ~v2_membered(X0)) [skolemisation 44,45]
51. v2_membered(sK0) [cnf transformation 41] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
52. v2_membered(sK1) [cnf transformation 41] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
53. ~r1_xxreal_0(k1_xxreal_2(k3_xboole_0(sK0,sK1)),k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1))) [cnf transformation 41] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. v1_xxreal_0(k1_xxreal_2(X0)) | ~v2_membered(X0) [cnf transformation 22] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. m1_xxreal_2(k3_xxreal_0(X2,X3),k3_xboole_0(X0,X1)) | ~m1_xxreal_2(X3,X1) | ~m1_xxreal_2(X2,X0) | ~v2_membered(X1) | ~v2_membered(X0) [cnf transformation 23] {a:0,w:17,wCS:6800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
56. m1_xxreal_2(X1,X0) | k1_xxreal_2(X0) != X1 | ~v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 46] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
57. r1_xxreal_0(X1,X3) | ~m1_xxreal_2(X3,X0) | k1_xxreal_2(X0) != X1 | ~v1_xxreal_0(X1) | ~v2_membered(X0) [cnf transformation 46] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
64. v2_membered(k3_xboole_0(X0,X1)) | ~v2_membered(X0) [cnf transformation 27] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
72. r1_xxreal_0(k1_xxreal_2(X0),X3) | ~m1_xxreal_2(X3,X0) | ~v1_xxreal_0(k1_xxreal_2(X0)) | ~v2_membered(X0) [equality resolution 57] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
73. m1_xxreal_2(k1_xxreal_2(X0),X0) | ~v1_xxreal_0(k1_xxreal_2(X0)) | ~v2_membered(X0) [equality resolution 56] {a:0,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
75. 1 <=> r1_xxreal_0(k1_xxreal_2(k3_xboole_0(sK0,sK1)),k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1))) [avatar definition]
77. ~r1_xxreal_0(k1_xxreal_2(k3_xboole_0(sK0,sK1)),k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1))) <- (~1) [avatar component clause 75] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
78. ~1 [avatar split clause 53,75]
80. 2 <=> v2_membered(sK1) [avatar definition]
82. v2_membered(sK1) <- (2) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
83. 2 [avatar split clause 52,80]
85. 3 <=> v2_membered(sK0) [avatar definition]
87. v2_membered(sK0) <- (3) [avatar component clause 85] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
88. 3 [avatar split clause 51,85]
115. m1_xxreal_2(k1_xxreal_2(X0),X0) | ~v2_membered(X0) [subsumption resolution 73,54] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
148. r1_xxreal_0(k1_xxreal_2(X0),X3) | ~m1_xxreal_2(X3,X0) | ~v2_membered(X0) [subsumption resolution 72,54] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
154. ~m1_xxreal_2(k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1)),k3_xboole_0(sK0,sK1)) | ~v2_membered(k3_xboole_0(sK0,sK1)) <- (~1) [resolution 148,77] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
156. 13 <=> v2_membered(k3_xboole_0(sK0,sK1)) [avatar definition]
158. ~v2_membered(k3_xboole_0(sK0,sK1)) <- (~13) [avatar component clause 156] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
160. 14 <=> m1_xxreal_2(k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1)),k3_xboole_0(sK0,sK1)) [avatar definition]
162. ~m1_xxreal_2(k3_xxreal_0(k1_xxreal_2(sK0),k1_xxreal_2(sK1)),k3_xboole_0(sK0,sK1)) <- (~14) [avatar component clause 160] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
163. ~13 | ~14 | 1 [avatar split clause 154,75,160,156]
193. ~v2_membered(sK0) <- (~13) [resolution 158,64] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
198. $false <- (3, ~13) [subsumption resolution 193,87] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
199. ~3 | 13 [avatar contradiction clause 198]
205. ~m1_xxreal_2(k1_xxreal_2(sK1),sK1) | ~m1_xxreal_2(k1_xxreal_2(sK0),sK0) | ~v2_membered(sK1) | ~v2_membered(sK0) <- (~14) [resolution 162,55] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
206. ~m1_xxreal_2(k1_xxreal_2(sK1),sK1) | ~m1_xxreal_2(k1_xxreal_2(sK0),sK0) | ~v2_membered(sK0) <- (2, ~14) [subsumption resolution 205,82] {a:2,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
207. ~m1_xxreal_2(k1_xxreal_2(sK1),sK1) | ~m1_xxreal_2(k1_xxreal_2(sK0),sK0) <- (2, 3, ~14) [subsumption resolution 206,87] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
209. 15 <=> m1_xxreal_2(k1_xxreal_2(sK0),sK0) [avatar definition]
211. ~m1_xxreal_2(k1_xxreal_2(sK0),sK0) <- (~15) [avatar component clause 209] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
213. 16 <=> m1_xxreal_2(k1_xxreal_2(sK1),sK1) [avatar definition]
215. ~m1_xxreal_2(k1_xxreal_2(sK1),sK1) <- (~16) [avatar component clause 213] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
216. ~15 | ~16 | ~2 | ~3 | 14 [avatar split clause 207,160,85,80,213,209]
218. ~v2_membered(sK0) <- (~15) [resolution 211,115] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
220. $false <- (3, ~15) [subsumption resolution 218,87] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
221. ~3 | 15 [avatar contradiction clause 220]
244. ~v2_membered(sK1) <- (~16) [resolution 215,115] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
246. $false <- (2, ~16) [subsumption resolution 244,82] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
247. ~2 | 16 [avatar contradiction clause 246]
248. $false [avatar sat refutation 78,83,88,163,199,216,221,247] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:180,thDist:-180}
% SZS output end Proof for xxreal_2__t67_xxreal_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 123
% Active clauses: 40
% Passive clauses: 53
% Final active clauses: 34
% Final passive clauses: 9
% Main loop iterations started: 44
% 
% >>> Simplifying Inferences
% Duplicate literals: 13
% Fw subsumption resolutions: 21
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 7
% Factoring: 1
% Forward superposition: 15
% Backward superposition: 6
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 15
% Sat splitting refutations: 14
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 27
% SAT solver unit clauses: 4
% SAT solver binary clauses: 16
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

