i: [1,110,0,2,1,1,0,0] l4_afvect01
a: [3,170,0,2,1,110]
s: 170
i: [1,109,0,2,1,1,0,0] l4_afvect01
a: [3,175,0,2,1,109]
s: 175
i: [1,108,0,2,1,1,0,0] l4_afvect01
a: [3,180,0,2,1,108]
s: 180
i: [1,112,0,4,1,1,0,0] l4_afvect01
a: [3,160,0,4,1,112]
s: 160
i: [1,111,0,4,1,1,0,0] l4_afvect01
a: [3,165,0,4,1,111]
s: 165
i: [1,113,0,6,1,1,0,0] l4_afvect01
a: [3,155,0,6,1,113]
s: 155
i: [1,114,0,2,1,0,0,255] spc1_boole
a: [3,185,0,2,1,114]
s: 185
i: [1,115,0,2,1,0,0,255] fc1_xboole_0
a: [3,190,0,2,1,115]
s: 190
i: [1,148,0,2,1,0,0,2] rc2_xboole_0
a: [3,195,0,2,1,148]
s: 195
i: [1,149,0,2,1,0,0,1] existence_l1_analoaf
a: [3,200,0,2,1,149]
s: 200
i: [1,150,0,2,1,0,0,2] existence_l1_struct_0
a: [3,205,0,2,1,150]
s: 205
i: [1,151,0,2,1,0,0,2] rc1_xboole_0
a: [3,210,0,2,1,151]
s: 210
i: [1,116,0,4,2,0,0,1] dt_l1_analoaf
s: 116
d: [2,212,1,2,1,1,40,116,170]
a: [3,217,1,2,1,212]
s: 217
i: [1,146,0,4,1,0,0,1] existence_m1_subset_1
s: 146
i: [1,131,0,5,2,0,0,255] t6_boole
s: 131
d: [2,213,1,2,1,1,40,116,200]
a: [3,222,1,2,1,213]
s: 222
i: [1,145,0,5,2,0,0,255] t7_boole
s: 145
i: [1,117,0,6,3,0,0,255] cc7_struct_0
s: 117
i: [1,118,0,6,3,0,0,255] cc5_struct_0
s: 118
i: [1,120,0,6,3,0,0,1] cc17_struct_0
s: 120
d: [2,226,1,4,2,1,40,117,205]
a: [3,232,1,2,1,226]
s: 232
i: [1,121,0,6,3,0,0,1] cc2_struct_0
s: 121
d: [2,228,2,4,2,1,40,117,222]
a: [3,241,2,2,1,228]
s: 241
i: [1,122,0,6,3,0,0,1] cc1_struct_0
s: 122
i: [1,125,0,6,3,0,0,255] cc4_struct_0
s: 125
i: [1,126,0,6,3,0,0,255] cc6_struct_0
s: 126
i: [1,142,0,6,2,0,0,255] t1_subset
s: 142
i: [1,147,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 147
i: [1,123,0,7,3,0,0,255] cc8_struct_0
s: 123
i: [1,127,0,7,3,0,0,255] cc11_struct_0
s: 127
i: [1,128,0,7,3,0,0,255] cc11_struct_0
s: 128
i: [1,129,0,7,3,0,0,255] cc9_struct_0
s: 129
i: [1,132,0,7,3,0,0,255] fc9_struct_0
s: 132
i: [1,136,0,7,3,0,0,1] fc2_struct_0
s: 136
i: [1,138,0,7,3,0,0,255] fc6_struct_0
s: 138
i: [1,139,0,7,3,0,0,255] fc8_struct_0
s: 139
i: [1,140,0,7,3,0,0,1] fc1_struct_0
s: 140
i: [1,141,0,7,3,0,0,255] fc7_struct_0
s: 141
i: [1,144,0,7,3,0,0,255] t8_boole
s: 144
i: [1,143,0,8,3,0,0,255] t2_subset
s: 143
d: [2,257,1,7,2,1,40,143,160]
a: [3,267,1,3,1,257]
s: 267
r: 267
i: [1,130,0,9,4,0,0,255] cc10_struct_0
s: 130
a: [3,266,1,3,1,257]
s: 266
i: [1,137,0,20,6,0,0,1] t1_afvect0
s: 137
i: [1,134,0,32,8,0,0,1] d4_diraf
s: 134
d: [2,274,2,4,2,1,40,267,136]
d: [2,275,2,2,1,2,21,274,217]
a: [3,278,2,2,1,275]
s: 278
i: [1,135,0,32,8,0,0,1] d4_diraf
s: 135
a: [3,263,1,4,1,257]
s: 263
i: [1,133,0,38,9,0,0,1] d4_diraf
s: 133
d: [2,258,1,7,2,1,40,143,165]
a: [3,272,1,4,1,258]
s: 272
d: [2,292,2,4,1,1,40,263,147]
a: [3,298,2,4,1,292]
s: 298
d: [2,259,1,6,2,0,40,143,146]
s: 259
d: [2,300,2,4,1,1,40,272,147]
a: [3,306,2,4,1,300]
s: 306
d: [2,224,1,7,3,0,44,131,131]
s: 224
d: [2,247,1,9,4,0,44,123,131]
s: 247
d: [2,308,2,6,2,0,40,259,147]
s: 308
d: [2,249,1,9,4,0,44,129,131]
s: 249
d: [2,288,1,40,12,0,40,134,137]
d: [2,289,1,22,7,0,18,288]
s: 289
d: [2,290,1,40,12,0,40,135,137]
d: [2,291,1,22,7,0,18,290]
s: 291
d: [2,325,3,0,0,7,52,180,175,278,170,160,165,155,291]
e: 325
d: [2,326,2,16,6,1,40,291,155]
d: [2,330,2,12,5,2,21,326,165]
d: [2,331,2,8,4,3,21,330,160]
d: [2,332,2,6,3,4,21,331,170]
d: [2,333,2,4,2,5,21,332,278]
d: [2,334,2,2,1,6,21,333,175]
d: [2,335,2,0,0,7,21,334,180]
e: 335
r: 266
r: 278
r: 263
r: 272
s: 267
a: [3,279,2,2,1,275]
s: 279
d: [2,345,3,0,0,3,52,217,180,279,120]
e: 345
d: [2,346,3,4,2,1,40,279,120]
d: [2,349,3,2,1,2,21,346,180]
d: [2,350,3,0,0,3,21,349,217]
e: 350
d: [2,339,2,4,2,5,21,332,175]
d: [2,340,2,2,1,6,21,339,180]
f: 350,340,212,108,109,110,111,112,113
% Refutation found. Thanks to Tanya!
% SZS status Theorem for afvect01__l4_afvect01
% SZS output start Proof for afvect01__l4_afvect01
1. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => r2_diraf(X0,X1,X2,X2,X1)))) [input l4_afvect01]
2. ~! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => r2_diraf(X0,X1,X2,X2,X1)))) [negated conjecture 1]
6. ! [X0] : (l1_struct_0(X0) => (~v7_struct_0(X0) => ~v2_struct_0(X0))) [input cc17_struct_0]
15. ! [X0] : ((l1_analoaf(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4)))))))) [input d4_diraf]
17. ! [X0] : (l1_analoaf(X0) => l1_struct_0(X0)) [input dt_l1_analoaf]
34. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => r2_analoaf(X0,X1,X2,X1,X2)))) [input t1_afvect0]
44. ? [X0] : (? [X1] : (? [X2] : (~r2_diraf(X0,X1,X2,X2,X1) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
45. ? [X0] : (? [X1] : (? [X2] : (~r2_diraf(X0,X1,X2,X2,X1) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) [flattening 44]
46. ! [X0] : (l1_struct_0(X0) | ~l1_analoaf(X0)) [ennf transformation 17]
51. ! [X0] : ((~v2_struct_0(X0) | v7_struct_0(X0)) | ~l1_struct_0(X0)) [ennf transformation 6]
52. ! [X0] : (~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0)) [flattening 51]
72. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | v2_struct_0(X0))) [ennf transformation 15]
73. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 72]
76. ! [X0] : (! [X1] : (! [X2] : (r2_analoaf(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 34]
77. ! [X0] : (! [X1] : (! [X2] : (r2_analoaf(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 76]
92. ? [X0] : (? [X1] : (? [X2] : (~r2_diraf(X0,X1,X2,X2,X1) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : (~r2_diraf(sK0,X1,X2,X2,X1) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0)) [choice axiom]
93. ? [X1] : (? [X2] : (~r2_diraf(sK0,X1,X2,X2,X1) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (~r2_diraf(sK0,sK1,X2,X2,sK1) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
94. ? [X2] : (~r2_diraf(sK0,sK1,X2,X2,sK1) & m1_subset_1(X2,u1_struct_0(sK0))) => (~r2_diraf(sK0,sK1,sK2,sK2,sK1) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
95. ((~r2_diraf(sK0,sK1,sK2,sK2,sK1) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0) [skolemisation 45,94,93,92]
96. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r2_diraf(X0,X1,X2,X3,X4) | (~r2_analoaf(X0,X1,X2,X4,X3) & ~r2_analoaf(X0,X1,X2,X3,X4))) & ((r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4)) | ~r2_diraf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [nnf transformation 73]
97. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r2_diraf(X0,X1,X2,X3,X4) | (~r2_analoaf(X0,X1,X2,X4,X3) & ~r2_analoaf(X0,X1,X2,X3,X4))) & (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4) | ~r2_diraf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 96]
108. ~v7_struct_0(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
109. v1_afvect0(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
110. l1_analoaf(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
111. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 95] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
112. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 95] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
113. ~r2_diraf(sK0,sK1,sK2,sK2,sK1) [cnf transformation 95] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
116. ~l1_analoaf(X0) | l1_struct_0(X0) [cnf transformation 46] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. ~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 52] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
135. ~r2_analoaf(X0,X1,X2,X4,X3) | r2_diraf(X0,X1,X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 97] {a:0,w:32,wCS:12800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
137. r2_analoaf(X0,X1,X2,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 77] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
153. 1 <=> r2_diraf(sK0,sK1,sK2,sK2,sK1) [avatar definition]
155. ~r2_diraf(sK0,sK1,sK2,sK2,sK1) <- (~1) [avatar component clause 153] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
156. ~1 [avatar split clause 113,153]
158. 2 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
160. m1_subset_1(sK2,u1_struct_0(sK0)) <- (2) [avatar component clause 158] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
161. 2 [avatar split clause 112,158]
163. 3 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
165. m1_subset_1(sK1,u1_struct_0(sK0)) <- (3) [avatar component clause 163] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
166. 3 [avatar split clause 111,163]
168. 4 <=> l1_analoaf(sK0) [avatar definition]
170. l1_analoaf(sK0) <- (4) [avatar component clause 168] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
171. 4 [avatar split clause 110,168]
173. 5 <=> v1_afvect0(sK0) [avatar definition]
175. v1_afvect0(sK0) <- (5) [avatar component clause 173] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
176. 5 [avatar split clause 109,173]
178. 6 <=> v7_struct_0(sK0) [avatar definition]
180. ~v7_struct_0(sK0) <- (~6) [avatar component clause 178] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
181. ~6 [avatar split clause 108,178]
212. l1_struct_0(sK0) <- (4) [resolution 116,170] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
215. 13 <=> l1_struct_0(sK0) [avatar definition]
217. l1_struct_0(sK0) <- (13) [avatar component clause 215] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
218. 13 | ~4 [avatar split clause 212,168,215]
277. 22 <=> v2_struct_0(sK0) [avatar definition]
279. v2_struct_0(sK0) <- (22) [avatar component clause 277] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
290. r2_diraf(X0,X1,X2,X2,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [resolution 135,137] {a:1,w:40,wCS:16000,sine:1,thAx:0,allAx:2,thDist:-2}
291. r2_diraf(X0,X1,X2,X2,X1) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [duplicate literal removal 290] {a:1,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
326. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~1) [resolution 291,155] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
330. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~1, 3) [subsumption resolution 326,165] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
331. ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~1, 2, 3) [subsumption resolution 330,160] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
332. v2_struct_0(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~1, 2, 3, 4) [subsumption resolution 331,170] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
339. v2_struct_0(sK0) | v7_struct_0(sK0) <- (~1, 2, 3, 4, 5) [subsumption resolution 332,175] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
340. v2_struct_0(sK0) <- (~1, 2, 3, 4, 5, ~6) [subsumption resolution 339,180] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
341. 22 | 1 | ~2 | ~3 | ~4 | ~5 | 6 [avatar split clause 340,178,173,168,163,158,153,277]
346. v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (22) [resolution 279,120] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
349. ~l1_struct_0(sK0) <- (~6, 22) [subsumption resolution 346,180] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
350. $false <- (~6, 13, 22) [subsumption resolution 349,217] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
351. 6 | ~13 | ~22 [avatar contradiction clause 350]
353. $false [avatar sat refutation 156,161,166,171,176,181,218,341,351] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:220,thDist:-220}
% SZS output end Proof for afvect01__l4_afvect01
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 44
% Generated clauses: 149
% Active clauses: 57
% Passive clauses: 64
% Final active clauses: 52
% Final passive clauses: 2
% Main loop iterations started: 62
% 
% >>> Simplifying Inferences
% Duplicate literals: 37
% Fw subsumption resolutions: 14
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 16
% 
% >>> Generating Inferences
% Binary resolution: 39
% Unit resulting resolution: 2
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 32
% SAT solver unit clauses: 12
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

