i: [1,43,0,12,2,1,0,0] t38_gate_1
a: [3,61,0,2,1,43]
s: 61
i: [1,42,0,12,2,1,0,0] t38_gate_1
a: [3,70,0,2,1,42]
s: 70
i: [1,41,0,12,2,1,0,0] t38_gate_1
a: [3,75,0,2,1,41]
s: 75
i: [1,40,0,12,2,1,0,0] t38_gate_1
a: [3,80,0,2,1,40]
s: 80
i: [1,39,0,12,2,1,0,0] t38_gate_1
a: [3,85,0,2,1,39]
s: 85
i: [1,38,0,12,2,1,0,0] t38_gate_1
a: [3,90,0,2,1,38]
s: 90
i: [1,37,0,12,2,1,0,0] t38_gate_1
a: [3,95,0,2,1,37]
s: 95
i: [1,36,0,12,2,1,0,0] t38_gate_1
a: [3,100,0,2,1,36]
s: 100
i: [1,44,0,2,1,0,0,255] fc1_xboole_0
a: [3,106,0,2,1,44]
s: 106
i: [1,56,0,2,1,0,0,1] rc2_xboole_0
a: [3,111,0,2,1,56]
s: 111
i: [1,57,0,2,1,0,0,1] rc1_xboole_0
a: [3,116,0,2,1,57]
s: 116
a: [3,64,0,10,1,43]
s: 64
i: [1,45,0,5,2,0,0,255] fc2_gate_1
s: 45
i: [1,46,0,5,2,0,0,255] fc1_gate_1
s: 46
i: [1,47,0,14,2,0,0,255] d32_gate_1
s: 47
i: [1,48,0,14,2,0,0,255] d32_gate_1
s: 48
i: [1,49,0,14,2,0,0,255] d32_gate_1
s: 49
i: [1,50,0,14,2,0,0,255] d32_gate_1
s: 50
i: [1,51,0,14,2,0,0,255] d32_gate_1
s: 51
i: [1,52,0,14,2,0,0,255] d32_gate_1
s: 52
i: [1,53,0,14,2,0,0,255] d32_gate_1
s: 53
i: [1,54,0,14,2,0,0,255] d32_gate_1
s: 54
i: [1,55,0,27,9,0,0,255] d32_gate_1
s: 55
d: [2,135,1,18,9,1,44,64,55]
d: [2,144,1,16,8,2,21,135,106]
d: [2,145,1,14,7,3,21,144,61]
d: [2,146,1,12,6,4,21,145,70]
d: [2,147,1,10,5,5,21,146,75]
d: [2,148,1,8,4,6,21,147,80]
d: [2,149,1,6,3,7,21,148,85]
d: [2,150,1,4,2,8,21,149,90]
d: [2,151,1,2,1,9,21,150,95]
d: [2,152,1,0,0,10,21,151,100]
e: 152
r: 64
r: 70
a: [3,69,0,2,1,42]
s: 69
a: [3,65,0,10,1,43]
s: 65
d: [2,160,1,5,2,1,44,65,53]
d: [2,163,1,3,1,2,21,160,69]
a: [3,167,1,3,1,163]
s: 167
d: [2,169,2,0,0,2,52,106,167,46]
e: 169
d: [2,170,2,2,1,1,40,167,46]
d: [2,172,2,0,0,2,21,170,106]
e: 172
r: 69
r: 75
r: 167
r: 80
r: 85
r: 90
r: 95
r: 100
r: 61
i: [1,26,0,26,9,1,0,0] t38_gate_1
d: [2,159,1,5,2,1,44,65,52]
d: [2,158,1,5,2,1,44,65,51]
d: [2,157,1,5,2,1,44,65,50]
d: [2,156,1,5,2,1,44,65,49]
d: [2,155,1,5,2,1,44,65,48]
d: [2,154,1,5,2,1,44,65,47]
d: [2,161,1,5,2,1,44,65,54]
f: 161,154,155,156,157,158,159,160,172,152,44,26,36,37,38,39,40,41,42,43
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t38_gate_1
% SZS output start Proof for gate_1__t38_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (~(v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)) & ~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & ~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)))) [input t38_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (~(v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)) & ~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & ~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (~(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) != k1_xboole_0 & ~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & (~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0))) [input d32_gate_1]
14. ~! [X0,X1,X2,X3,X4,X5,X6,X7] : (~(v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)) & ~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & ~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)))) [flattening 2]
15. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) != k1_xboole_0 & ~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) & (~(~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0))) [flattening 13]
16. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)) & (v1_xboole_0(X7) | v1_xboole_0(X6) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) | (~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | v1_xboole_0(X7) | v1_xboole_0(X6) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) & (k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | (~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)))) [ennf transformation 15]
20. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)) & (v1_xboole_0(X7) | v1_xboole_0(X6) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) | (~v1_xboole_0(X7) & ~v1_xboole_0(X6) & ~v1_xboole_0(X5) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7)))) => ((v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) & (v1_xboole_0(sK7) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0))) | (~v1_xboole_0(sK7) & ~v1_xboole_0(sK6) & ~v1_xboole_0(sK5) & ~v1_xboole_0(sK4) & ~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0) & ~v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)))) [choice axiom]
21. (v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) & (v1_xboole_0(sK7) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0))) | (~v1_xboole_0(sK7) & ~v1_xboole_0(sK6) & ~v1_xboole_0(sK5) & ~v1_xboole_0(sK4) & ~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0) & ~v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7))) [skolemisation 16,20]
26. v1_xboole_0(sK7) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | ~v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [cnf transformation 21] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) | ~v1_xboole_0(sK0) [cnf transformation 21] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
37. v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) | ~v1_xboole_0(sK1) [cnf transformation 21] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
38. v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) | ~v1_xboole_0(sK2) [cnf transformation 21] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
39. v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) | ~v1_xboole_0(sK3) [cnf transformation 21] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
40. v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) | ~v1_xboole_0(sK4) [cnf transformation 21] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
41. v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) | ~v1_xboole_0(sK5) [cnf transformation 21] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) | ~v1_xboole_0(sK6) [cnf transformation 21] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) | ~v1_xboole_0(sK7) [cnf transformation 21] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
44. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
46. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
47. k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
48. k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X1) [cnf transformation 19] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
49. k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X2) [cnf transformation 19] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
50. k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X3) [cnf transformation 19] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
51. k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X4) [cnf transformation 19] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
52. k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X5) [cnf transformation 19] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
53. k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X6) [cnf transformation 19] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
54. k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_gate_1(k1_xboole_0) | ~v1_xboole_0(X7) [cnf transformation 19] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
55. k32_gate_1(X0,X1,X2,X3,X4,X5,X6,X7) = k1_xboole_0 | v1_xboole_0(X7) | v1_xboole_0(X6) | v1_xboole_0(X5) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:27,wCS:10800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
59. 1 <=> v1_xboole_0(sK7) [avatar definition]
61. ~v1_xboole_0(sK7) <- (~1) [avatar component clause 59] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
63. 2 <=> v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) [avatar definition]
64. ~v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) <- (~2) [avatar component clause 63] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
65. v1_xboole_0(k32_gate_1(sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7)) <- (2) [avatar component clause 63] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66. ~1 | 2 [avatar split clause 43,63,59]
68. 3 <=> v1_xboole_0(sK6) [avatar definition]
70. ~v1_xboole_0(sK6) <- (~3) [avatar component clause 68] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. ~3 | 2 [avatar split clause 42,63,68]
73. 4 <=> v1_xboole_0(sK5) [avatar definition]
75. ~v1_xboole_0(sK5) <- (~4) [avatar component clause 73] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. ~4 | 2 [avatar split clause 41,63,73]
78. 5 <=> v1_xboole_0(sK4) [avatar definition]
80. ~v1_xboole_0(sK4) <- (~5) [avatar component clause 78] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
81. ~5 | 2 [avatar split clause 40,63,78]
83. 6 <=> v1_xboole_0(sK3) [avatar definition]
85. ~v1_xboole_0(sK3) <- (~6) [avatar component clause 83] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
86. ~6 | 2 [avatar split clause 39,63,83]
88. 7 <=> v1_xboole_0(sK2) [avatar definition]
90. ~v1_xboole_0(sK2) <- (~7) [avatar component clause 88] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
91. ~7 | 2 [avatar split clause 38,63,88]
93. 8 <=> v1_xboole_0(sK1) [avatar definition]
95. ~v1_xboole_0(sK1) <- (~8) [avatar component clause 93] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
96. ~8 | 2 [avatar split clause 37,63,93]
98. 9 <=> v1_xboole_0(sK0) [avatar definition]
100. ~v1_xboole_0(sK0) <- (~9) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
101. ~9 | 2 [avatar split clause 36,63,98]
102. ~2 | 9 | 8 | 7 | 6 | 5 | 4 | 3 | 1 [avatar split clause 26,59,68,73,78,83,88,93,98,63]
104. 10 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
106. v1_xboole_0(k1_xboole_0) <- (10) [avatar component clause 104] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
107. 10 [avatar split clause 44,104]
135. ~v1_xboole_0(k1_xboole_0) | v1_xboole_0(sK7) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~2) [superposition 64,55] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
144. v1_xboole_0(sK7) | v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~2, 10) [subsumption resolution 135,106] {a:1,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
145. v1_xboole_0(sK6) | v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1, ~2, 10) [subsumption resolution 144,61] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
146. v1_xboole_0(sK5) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1, ~2, ~3, 10) [subsumption resolution 145,70] {a:1,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
147. v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1, ~2, ~3, ~4, 10) [subsumption resolution 146,75] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
148. v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1, ~2, ~3, ~4, ~5, 10) [subsumption resolution 147,80] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
149. v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1, ~2, ~3, ~4, ~5, ~6, 10) [subsumption resolution 148,85] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
150. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~1, ~2, ~3, ~4, ~5, ~6, ~7, 10) [subsumption resolution 149,90] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
151. v1_xboole_0(sK0) <- (~1, ~2, ~3, ~4, ~5, ~6, ~7, ~8, 10) [subsumption resolution 150,95] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
152. $false <- (~1, ~2, ~3, ~4, ~5, ~6, ~7, ~8, ~9, 10) [subsumption resolution 151,100] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
153. 1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | 9 | ~10 [avatar contradiction clause 152]
154. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK0) <- (2) [superposition 65,47] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
155. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK1) <- (2) [superposition 65,48] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
156. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK2) <- (2) [superposition 65,49] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
157. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK3) <- (2) [superposition 65,50] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
158. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK4) <- (2) [superposition 65,51] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
159. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK5) <- (2) [superposition 65,52] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
160. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK6) <- (2) [superposition 65,53] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
161. v1_xboole_0(k1_gate_1(k1_xboole_0)) | ~v1_xboole_0(sK7) <- (2) [superposition 65,54] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
165. 13 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
167. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (13) [avatar component clause 165] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
170. ~v1_xboole_0(k1_xboole_0) <- (13) [resolution 167,46] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
172. $false <- (10, 13) [subsumption resolution 170,106] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
173. ~10 | ~13 [avatar contradiction clause 172]
174. ~3 | 13 | ~2 [avatar split clause 160,63,165,68]
175. ~4 | 13 | ~2 [avatar split clause 159,63,165,73]
176. ~5 | 13 | ~2 [avatar split clause 158,63,165,78]
177. ~6 | 13 | ~2 [avatar split clause 157,63,165,83]
178. ~7 | 13 | ~2 [avatar split clause 156,63,165,88]
179. ~8 | 13 | ~2 [avatar split clause 155,63,165,93]
180. ~9 | 13 | ~2 [avatar split clause 154,63,165,98]
181. ~1 | 13 | ~2 [avatar split clause 161,63,165,59]
182. $false [avatar sat refutation 66,71,76,81,86,91,96,101,102,107,153,173,174,175,176,177,178,179,180,181] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:117,thDist:-117}
% SZS output end Proof for gate_1__t38_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 10
% 
% >>> Saturation
% Initial clauses: 32
% Generated clauses: 118
% Active clauses: 26
% Passive clauses: 41
% Final active clauses: 15
% Final passive clauses: 9
% Main loop iterations started: 26
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 26
% Forward subsumptions: 16
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 17
% Backward superposition: 17
% 
% >>> AVATAR
% Split clauses: 17
% Split components: 41
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 26
% SAT solver unit clauses: 3
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

