i: [1,170,0,12,3,1,0,0] t1_finsub_1
a: [3,252,0,2,1,170]
s: 252
i: [1,171,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,283,0,2,1,171]
s: 283
i: [1,172,0,2,1,0,0,3] fc1_xboole_0
a: [3,288,0,2,1,172]
s: 288
i: [1,242,0,2,1,0,0,2] rc2_xboole_0
a: [3,298,0,2,1,242]
s: 298
i: [1,243,0,2,1,0,0,2] rc1_xboole_0
a: [3,303,0,2,1,243]
s: 303
i: [1,245,0,2,1,0,0,2] rc1_finset_1
a: [3,308,0,2,1,245]
s: 308
i: [1,244,0,2,1,0,0,2] rc1_finset_1
a: [3,313,0,2,1,244]
s: 313
i: [1,248,0,2,1,0,0,255] rc8_finset_1
a: [3,318,0,2,1,248]
s: 318
i: [1,247,0,2,1,0,0,255] rc8_finset_1
a: [3,323,0,2,1,247]
s: 323
i: [1,246,0,2,1,0,0,255] rc8_finset_1
a: [3,328,0,2,1,246]
s: 328
i: [1,167,0,13,4,1,0,0] t1_finsub_1
a: [3,274,0,11,3,167]
s: 274
i: [1,166,0,13,4,1,0,0] t1_finsub_1
a: [3,278,0,11,3,166]
s: 278
i: [1,174,0,3,1,0,0,2] fc1_subset_1
s: 174
i: [1,228,0,3,1,0,0,2] rc2_subset_1
s: 228
i: [1,231,0,3,1,0,0,255] reflexivity_r1_tarski
s: 231
i: [1,173,0,3,1,0,0,255] d2_xboole_0
a: [3,293,0,3,1,173]
s: 293
r: 283
i: [1,183,0,4,2,0,0,255] cc5_finset_1
s: 183
i: [1,196,0,4,2,0,0,1] cc1_finsub_1
s: 196
i: [1,197,0,4,2,0,0,1] cc1_finsub_1
s: 197
i: [1,198,0,4,2,0,0,255] cc4_finset_1
s: 198
i: [1,205,0,4,2,0,0,255] cc6_finset_1
s: 205
d: [2,331,1,2,1,1,40,205,288]
a: [3,337,1,2,1,331]
s: 337
i: [1,206,0,4,2,0,0,2] cc1_finset_1
s: 206
d: [2,333,1,2,1,1,40,205,303]
a: [3,342,1,2,1,333]
s: 342
i: [1,226,0,4,1,0,0,2] existence_m1_subset_1
s: 226
i: [1,230,0,4,1,0,0,255] rc3_subset_1
s: 230
d: [2,344,1,2,1,1,40,206,288]
a: [3,350,1,2,1,344]
s: 350
i: [1,175,0,5,1,0,0,2] t4_boole
s: 175
d: [2,346,1,2,1,1,40,206,303]
a: [3,355,1,2,1,346]
s: 355
i: [1,176,0,5,1,0,0,2] t3_boole
s: 176
i: [1,177,0,5,1,0,0,3] t1_boole
s: 177
d: [2,332,1,3,1,0,40,205,228]
s: 332
i: [1,179,0,5,2,0,0,255] rc6_subset_1
s: 179
d: [2,345,1,3,1,0,40,206,228]
s: 345
i: [1,181,0,5,2,0,0,255] rc7_finset_1
s: 181
i: [1,182,0,5,2,0,0,255] rc7_finset_1
s: 182
i: [1,187,0,5,2,0,0,2] rc1_subset_1
s: 187
i: [1,189,0,5,2,0,0,255] rc5_subset_1
s: 189
i: [1,190,0,5,2,0,0,255] rc5_subset_1
s: 190
i: [1,192,0,5,2,0,0,2] rc2_finset_1
s: 192
i: [1,193,0,5,2,0,0,2] rc2_finset_1
s: 193
i: [1,202,0,5,2,0,0,255] fc31_finset_1
s: 202
i: [1,203,0,5,2,0,0,2] fc17_finset_1
s: 203
i: [1,207,0,5,2,0,0,3] t6_boole
s: 207
i: [1,223,0,5,2,0,0,1] t7_boole
s: 223
i: [1,227,0,5,1,0,0,2] rc2_subset_1
s: 227
i: [1,229,0,5,1,0,0,255] rc3_subset_1
s: 229
i: [1,232,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 232
d: [2,358,1,5,2,0,40,190,183]
s: 358
i: [1,195,0,6,2,0,0,255] rc4_subset_1
s: 195
i: [1,210,0,6,3,0,0,1] cc2_finsub_1
s: 210
i: [1,212,0,6,2,0,0,2] d3_finsub_1
s: 212
i: [1,213,0,6,2,0,0,2] d3_finsub_1
s: 213
d: [2,370,1,4,2,0,40,212,223]
s: 370
i: [1,216,0,6,2,0,0,2] d1_finsub_1
s: 216
i: [1,217,0,6,2,0,0,2] d1_finsub_1
s: 217
d: [2,372,2,2,1,1,40,370,288]
a: [3,378,2,2,1,372]
s: 378
i: [1,219,0,6,2,0,0,1] t1_subset
s: 219
d: [2,374,2,2,1,1,40,370,303]
a: [3,383,2,2,1,374]
s: 383
i: [1,233,0,6,1,0,0,1] dt_k6_subset_1
s: 233
i: [1,236,0,6,2,0,0,1] fc5_xboole_0
s: 236
d: [2,387,3,4,2,1,40,378,210]
a: [3,392,3,2,1,387]
s: 392
i: [1,237,0,6,2,0,0,1] fc4_xboole_0
s: 237
d: [2,404,3,4,2,1,40,383,210]
a: [3,408,3,2,1,404]
s: 408
i: [1,238,0,6,2,0,0,2] fc12_finset_1
s: 238
i: [1,239,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 239
d: [2,398,1,11,3,1,40,219,274]
s: 398
i: [1,178,0,7,2,0,0,255] rc6_subset_1
s: 178
d: [2,399,1,11,3,1,40,219,278]
s: 399
i: [1,180,0,7,2,0,0,255] rc7_finset_1
s: 180
i: [1,186,0,7,2,0,0,2] rc1_subset_1
s: 186
d: [2,421,1,11,3,1,40,239,274]
s: 421
i: [1,188,0,7,2,0,0,255] rc5_subset_1
s: 188
d: [2,422,1,11,3,1,40,239,278]
s: 422
i: [1,191,0,7,2,0,0,2] rc2_finset_1
s: 191
i: [1,194,0,7,2,0,0,255] rc4_subset_1
s: 194
d: [2,373,2,3,1,0,40,370,228]
s: 373
i: [1,200,0,7,3,0,0,255] cc7_finset_1
s: 200
d: [2,434,2,12,4,1,40,200,399]
a: [3,449,2,2,1,434]
s: 449
i: [1,221,0,7,2,0,0,255] t3_subset
s: 221
i: [1,222,0,7,3,0,0,2] t8_boole
s: 222
d: [2,371,1,4,2,0,40,213,223]
s: 371
i: [1,234,0,7,1,0,0,1] redefinition_k6_subset_1
s: 234
r: 238
d: [2,385,1,4,2,0,40,216,223]
s: 385
d: [2,464,4,0,0,2,52,392,288,385]
e: 464
d: [2,465,4,0,0,2,52,408,303,385]
e: 465
d: [2,466,2,2,1,1,40,385,288]
d: [2,471,2,0,0,2,21,466,392]
e: 471
d: [2,468,2,2,1,1,40,385,303]
d: [2,473,2,0,0,2,21,468,408]
e: 473
r: 392
r: 408
d: [2,459,0,6,2,0,23,238,234]
s: 459
i: [1,235,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 235
a: [3,391,3,2,1,387]
s: 391
i: [1,199,0,8,3,0,0,255] cc5_subset_1
s: 199
a: [3,396,3,2,1,387]
s: 396
i: [1,201,0,8,3,0,0,255] cc8_finset_1
s: 201
i: [1,204,0,8,3,0,0,2] cc2_finset_1
s: 204
a: [3,407,3,2,1,404]
s: 407
i: [1,208,0,8,3,0,0,2] cc1_subset_1
s: 208
a: [3,412,3,2,1,404]
s: 412
i: [1,220,0,8,3,0,0,1] t2_subset
s: 220
i: [1,240,0,8,3,0,0,255] fc33_finset_1
s: 240
d: [2,467,2,3,1,0,40,385,228]
s: 467
i: [1,241,0,8,3,0,0,1] fc9_finset_1
s: 241
d: [2,386,1,4,2,0,40,217,223]
s: 386
i: [1,209,0,9,3,0,0,255] cc4_subset_1
s: 209
i: [1,214,0,9,2,0,0,2] d3_finsub_1
s: 214
r: 252
r: 421
r: 398
r: 274
a: [3,251,0,2,1,170]
s: 251
r: 251
s: 252
d: [2,585,1,10,3,1,40,214,274]
d: [2,586,1,6,2,1,21,585,212]
d: [2,587,1,2,1,1,21,586,213]
a: [3,591,1,2,1,587]
s: 591
i: [1,169,0,5,2,1,0,0] t1_finsub_1
a: [3,265,0,3,1,169]
s: 265
r: 265
a: [3,264,0,3,1,169]
s: 264
d: [2,594,1,2,1,1,40,251,196]
a: [3,598,1,2,1,594]
s: 598
i: [1,168,0,5,2,1,0,0] t1_finsub_1
a: [3,270,0,3,1,168]
s: 270
d: [2,606,1,2,1,1,40,265,223]
a: [3,620,1,2,1,606]
s: 620
a: [3,260,0,5,1,170]
s: 260
s: 274
d: [2,622,1,3,1,1,40,270,239]
a: [3,628,1,3,1,622]
s: 628
i: [1,218,0,9,2,0,0,2] d1_finsub_1
s: 218
d: [2,623,1,3,1,1,40,270,219]
a: [3,633,1,3,1,623]
s: 633
i: [1,225,0,9,3,0,0,1] t5_subset
s: 225
i: [1,224,0,10,3,0,0,1] t4_subset
s: 224
d: [2,641,1,10,3,1,40,218,278]
d: [2,642,1,8,2,2,21,641,598]
a: [3,646,1,4,1,642]
s: 646
d: [2,676,2,0,0,2,52,598,646,217]
e: 676
d: [2,677,2,2,1,1,40,646,217]
d: [2,679,2,0,0,2,21,677,598]
e: 679
r: 252
r: 264
r: 598
r: 260
s: 251
a: [3,599,1,2,1,594]
s: 599
s: 265
a: [3,256,0,5,1,170]
s: 256
d: [2,693,1,0,0,4,52,265,270,256,278]
e: 693
d: [2,694,1,6,2,2,40,256,278]
d: [2,696,1,3,1,3,21,694,270]
d: [2,697,1,0,0,4,21,696,265]
e: 697
r: 422
r: 399
r: 278
d: [2,604,1,3,1,1,40,265,239]
a: [3,610,1,3,1,604]
s: 610
a: [3,259,0,5,1,170]
s: 259
d: [2,605,1,3,1,1,40,265,219]
a: [3,615,1,3,1,605]
s: 615
i: [1,184,0,11,4,0,0,255] cc3_subset_1
s: 184
i: [1,185,0,11,4,0,0,255] cc2_subset_1
s: 185
d: [2,699,1,5,1,1,40,259,239]
a: [3,705,1,5,1,699]
s: 705
i: [1,211,0,13,4,0,0,2] d3_finsub_1
s: 211
d: [2,700,1,5,1,1,40,259,219]
a: [3,710,1,5,1,700]
s: 710
i: [1,215,0,13,4,0,0,2] d1_finsub_1
s: 215
d: [2,737,2,0,0,4,52,599,270,265,256,215]
e: 737
d: [2,739,1,8,3,1,40,215,256]
d: [2,749,1,5,2,2,21,739,265]
d: [2,750,1,2,1,3,21,749,270]
d: [2,751,1,0,0,4,21,750,599]
e: 751
r: 256
r: 259
r: 274
r: 710
s: 260
d: [2,755,2,0,0,4,52,591,270,265,260,211]
e: 755
d: [2,756,1,8,3,1,40,260,211]
d: [2,758,1,5,2,2,21,756,265]
d: [2,759,1,2,1,3,21,758,270]
d: [2,760,1,0,0,4,21,759,591]
e: 760
r: 251
r: 260
r: 599
r: 646
s: 252
s: 259
s: 598
s: 274
s: 278
a: [3,645,1,4,1,642]
s: 645
a: [3,650,1,4,1,642]
s: 650
d: [2,794,2,0,0,2,52,598,650,216]
e: 794
d: [2,795,2,2,1,1,40,650,216]
d: [2,797,2,0,0,2,21,795,598]
e: 797
d: [2,593,1,2,1,1,40,251,197]
d: [2,601,2,4,2,1,40,591,210]
d: [2,686,1,2,1,1,40,251,196]
d: [2,770,1,10,3,1,40,278,218]
d: [2,779,1,8,2,2,21,770,598]
f: 797,779,760,751,686,677,601,593,587,166,167,168,169,170
% Refutation found. Thanks to Tanya!
% SZS status Theorem for finsub_1__t1_finsub_1
% SZS output start Proof for finsub_1__t1_finsub_1
1. ! [X0] : (v4_finsub_1(X0) <=> ! [X1] : ! [X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (r2_hidden(k6_subset_1(X1,X2),X0) & r2_hidden(k2_xboole_0(X1,X2),X0)))) [input t1_finsub_1]
2. ~! [X0] : (v4_finsub_1(X0) <=> ! [X1] : ! [X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (r2_hidden(k6_subset_1(X1,X2),X0) & r2_hidden(k2_xboole_0(X1,X2),X0)))) [negated conjecture 1]
5. ! [X0] : (v4_finsub_1(X0) => (v3_finsub_1(X0) & v1_finsub_1(X0))) [input cc1_finsub_1]
8. ! [X0] : ((v3_finsub_1(X0) & v1_finsub_1(X0)) => v4_finsub_1(X0)) [input cc2_finsub_1]
19. ! [X0] : (v1_finsub_1(X0) <=> ! [X1] : ! [X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => r2_hidden(k2_xboole_0(X1,X2),X0))) [input d1_finsub_1]
21. ! [X0] : (v3_finsub_1(X0) <=> ! [X1] : ! [X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => r2_hidden(k6_subset_1(X1,X2),X0))) [input d3_finsub_1]
65. ~! [X0] : (v4_finsub_1(X0) <=> ! [X1,X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (r2_hidden(k6_subset_1(X1,X2),X0) & r2_hidden(k2_xboole_0(X1,X2),X0)))) [flattening 2]
66. ! [X0] : (v3_finsub_1(X0) <=> ! [X1,X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => r2_hidden(k6_subset_1(X1,X2),X0))) [flattening 21]
67. ! [X0] : (v1_finsub_1(X0) <=> ! [X1,X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => r2_hidden(k2_xboole_0(X1,X2),X0))) [flattening 19]
78. ? [X0] : (v4_finsub_1(X0) <~> ! [X1,X2] : ((r2_hidden(k6_subset_1(X1,X2),X0) & r2_hidden(k2_xboole_0(X1,X2),X0)) | (~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)))) [ennf transformation 65]
79. ? [X0] : (v4_finsub_1(X0) <~> ! [X1,X2] : ((r2_hidden(k6_subset_1(X1,X2),X0) & r2_hidden(k2_xboole_0(X1,X2),X0)) | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0))) [flattening 78]
91. ! [X0] : ((v3_finsub_1(X0) & v1_finsub_1(X0)) | ~v4_finsub_1(X0)) [ennf transformation 5]
104. ! [X0] : (v4_finsub_1(X0) | (~v3_finsub_1(X0) | ~v1_finsub_1(X0))) [ennf transformation 8]
105. ! [X0] : (v4_finsub_1(X0) | ~v3_finsub_1(X0) | ~v1_finsub_1(X0)) [flattening 104]
106. ! [X0] : (v3_finsub_1(X0) <=> ! [X1,X2] : (r2_hidden(k6_subset_1(X1,X2),X0) | (~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)))) [ennf transformation 66]
107. ! [X0] : (v3_finsub_1(X0) <=> ! [X1,X2] : (r2_hidden(k6_subset_1(X1,X2),X0) | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0))) [flattening 106]
108. ! [X0] : (v1_finsub_1(X0) <=> ! [X1,X2] : (r2_hidden(k2_xboole_0(X1,X2),X0) | (~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)))) [ennf transformation 67]
109. ! [X0] : (v1_finsub_1(X0) <=> ! [X1,X2] : (r2_hidden(k2_xboole_0(X1,X2),X0) | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0))) [flattening 108]
127. ? [X0] : ((? [X1,X2] : ((~r2_hidden(k6_subset_1(X1,X2),X0) | ~r2_hidden(k2_xboole_0(X1,X2),X0)) & r2_hidden(X2,X0) & r2_hidden(X1,X0)) | ~v4_finsub_1(X0)) & (! [X1,X2] : ((r2_hidden(k6_subset_1(X1,X2),X0) & r2_hidden(k2_xboole_0(X1,X2),X0)) | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)) | v4_finsub_1(X0))) [nnf transformation 79]
128. ? [X0] : ((? [X1,X2] : ((~r2_hidden(k6_subset_1(X1,X2),X0) | ~r2_hidden(k2_xboole_0(X1,X2),X0)) & r2_hidden(X2,X0) & r2_hidden(X1,X0)) | ~v4_finsub_1(X0)) & (! [X3,X4] : ((r2_hidden(k6_subset_1(X3,X4),X0) & r2_hidden(k2_xboole_0(X3,X4),X0)) | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | v4_finsub_1(X0))) [rectify 127]
129. ? [X0] : ((? [X1,X2] : ((~r2_hidden(k6_subset_1(X1,X2),X0) | ~r2_hidden(k2_xboole_0(X1,X2),X0)) & r2_hidden(X2,X0) & r2_hidden(X1,X0)) | ~v4_finsub_1(X0)) & (! [X3,X4] : ((r2_hidden(k6_subset_1(X3,X4),X0) & r2_hidden(k2_xboole_0(X3,X4),X0)) | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | v4_finsub_1(X0))) => ((? [X2,X1] : ((~r2_hidden(k6_subset_1(X1,X2),sK0) | ~r2_hidden(k2_xboole_0(X1,X2),sK0)) & r2_hidden(X2,sK0) & r2_hidden(X1,sK0)) | ~v4_finsub_1(sK0)) & (! [X4,X3] : ((r2_hidden(k6_subset_1(X3,X4),sK0) & r2_hidden(k2_xboole_0(X3,X4),sK0)) | ~r2_hidden(X4,sK0) | ~r2_hidden(X3,sK0)) | v4_finsub_1(sK0))) [choice axiom]
130. ? [X2,X1] : ((~r2_hidden(k6_subset_1(X1,X2),sK0) | ~r2_hidden(k2_xboole_0(X1,X2),sK0)) & r2_hidden(X2,sK0) & r2_hidden(X1,sK0)) => ((~r2_hidden(k6_subset_1(sK1,sK2),sK0) | ~r2_hidden(k2_xboole_0(sK1,sK2),sK0)) & r2_hidden(sK2,sK0) & r2_hidden(sK1,sK0)) [choice axiom]
131. (((~r2_hidden(k6_subset_1(sK1,sK2),sK0) | ~r2_hidden(k2_xboole_0(sK1,sK2),sK0)) & r2_hidden(sK2,sK0) & r2_hidden(sK1,sK0)) | ~v4_finsub_1(sK0)) & (! [X3,X4] : ((r2_hidden(k6_subset_1(X3,X4),sK0) & r2_hidden(k2_xboole_0(X3,X4),sK0)) | ~r2_hidden(X4,sK0) | ~r2_hidden(X3,sK0)) | v4_finsub_1(sK0)) [skolemisation 128,130,129]
144. ! [X0] : ((v3_finsub_1(X0) | ? [X1,X2] : (~r2_hidden(k6_subset_1(X1,X2),X0) & r2_hidden(X2,X0) & r2_hidden(X1,X0))) & (! [X1,X2] : (r2_hidden(k6_subset_1(X1,X2),X0) | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)) | ~v3_finsub_1(X0))) [nnf transformation 107]
145. ! [X0] : ((v3_finsub_1(X0) | ? [X1,X2] : (~r2_hidden(k6_subset_1(X1,X2),X0) & r2_hidden(X2,X0) & r2_hidden(X1,X0))) & (! [X3,X4] : (r2_hidden(k6_subset_1(X3,X4),X0) | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | ~v3_finsub_1(X0))) [rectify 144]
146. ! [X0] : (? [X1,X2] : (~r2_hidden(k6_subset_1(X1,X2),X0) & r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (~r2_hidden(k6_subset_1(sK9(X0),sK10(X0)),X0) & r2_hidden(sK10(X0),X0) & r2_hidden(sK9(X0),X0))) [choice axiom]
147. ! [X0] : ((v3_finsub_1(X0) | (~r2_hidden(k6_subset_1(sK9(X0),sK10(X0)),X0) & r2_hidden(sK10(X0),X0) & r2_hidden(sK9(X0),X0))) & (! [X3,X4] : (r2_hidden(k6_subset_1(X3,X4),X0) | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | ~v3_finsub_1(X0))) [skolemisation 145,146]
148. ! [X0] : ((v1_finsub_1(X0) | ? [X1,X2] : (~r2_hidden(k2_xboole_0(X1,X2),X0) & r2_hidden(X2,X0) & r2_hidden(X1,X0))) & (! [X1,X2] : (r2_hidden(k2_xboole_0(X1,X2),X0) | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)) | ~v1_finsub_1(X0))) [nnf transformation 109]
149. ! [X0] : ((v1_finsub_1(X0) | ? [X1,X2] : (~r2_hidden(k2_xboole_0(X1,X2),X0) & r2_hidden(X2,X0) & r2_hidden(X1,X0))) & (! [X3,X4] : (r2_hidden(k2_xboole_0(X3,X4),X0) | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | ~v1_finsub_1(X0))) [rectify 148]
150. ! [X0] : (? [X1,X2] : (~r2_hidden(k2_xboole_0(X1,X2),X0) & r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (~r2_hidden(k2_xboole_0(sK11(X0),sK12(X0)),X0) & r2_hidden(sK12(X0),X0) & r2_hidden(sK11(X0),X0))) [choice axiom]
151. ! [X0] : ((v1_finsub_1(X0) | (~r2_hidden(k2_xboole_0(sK11(X0),sK12(X0)),X0) & r2_hidden(sK12(X0),X0) & r2_hidden(sK11(X0),X0))) & (! [X3,X4] : (r2_hidden(k2_xboole_0(X3,X4),X0) | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | ~v1_finsub_1(X0))) [skolemisation 149,150]
166. r2_hidden(k2_xboole_0(X3,X4),sK0) | ~r2_hidden(X4,sK0) | ~r2_hidden(X3,sK0) | v4_finsub_1(sK0) [cnf transformation 131] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
167. r2_hidden(k6_subset_1(X3,X4),sK0) | ~r2_hidden(X4,sK0) | ~r2_hidden(X3,sK0) | v4_finsub_1(sK0) [cnf transformation 131] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
168. r2_hidden(sK1,sK0) | ~v4_finsub_1(sK0) [cnf transformation 131] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
169. r2_hidden(sK2,sK0) | ~v4_finsub_1(sK0) [cnf transformation 131] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
170. ~r2_hidden(k6_subset_1(sK1,sK2),sK0) | ~r2_hidden(k2_xboole_0(sK1,sK2),sK0) | ~v4_finsub_1(sK0) [cnf transformation 131] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
196. ~v4_finsub_1(X0) | v1_finsub_1(X0) [cnf transformation 91] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
197. ~v4_finsub_1(X0) | v3_finsub_1(X0) [cnf transformation 91] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
210. ~v3_finsub_1(X0) | v4_finsub_1(X0) | ~v1_finsub_1(X0) [cnf transformation 105] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
211. r2_hidden(k6_subset_1(X3,X4),X0) | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0) | ~v3_finsub_1(X0) [cnf transformation 147] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
212. r2_hidden(sK9(X0),X0) | v3_finsub_1(X0) [cnf transformation 147] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
213. r2_hidden(sK10(X0),X0) | v3_finsub_1(X0) [cnf transformation 147] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
214. ~r2_hidden(k6_subset_1(sK9(X0),sK10(X0)),X0) | v3_finsub_1(X0) [cnf transformation 147] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
215. r2_hidden(k2_xboole_0(X3,X4),X0) | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0) | ~v1_finsub_1(X0) [cnf transformation 151] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
216. r2_hidden(sK11(X0),X0) | v1_finsub_1(X0) [cnf transformation 151] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
217. r2_hidden(sK12(X0),X0) | v1_finsub_1(X0) [cnf transformation 151] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
218. ~r2_hidden(k2_xboole_0(sK11(X0),sK12(X0)),X0) | v1_finsub_1(X0) [cnf transformation 151] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
250. 1 <=> v4_finsub_1(sK0) [avatar definition]
251. v4_finsub_1(sK0) <- (1) [avatar component clause 250] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
254. 2 <=> r2_hidden(k2_xboole_0(sK1,sK2),sK0) [avatar definition]
256. ~r2_hidden(k2_xboole_0(sK1,sK2),sK0) <- (~2) [avatar component clause 254] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
258. 3 <=> r2_hidden(k6_subset_1(sK1,sK2),sK0) [avatar definition]
260. ~r2_hidden(k6_subset_1(sK1,sK2),sK0) <- (~3) [avatar component clause 258] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
261. ~1 | ~2 | ~3 [avatar split clause 170,258,254,250]
263. 4 <=> r2_hidden(sK2,sK0) [avatar definition]
265. r2_hidden(sK2,sK0) <- (4) [avatar component clause 263] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
266. ~1 | 4 [avatar split clause 169,263,250]
268. 5 <=> r2_hidden(sK1,sK0) [avatar definition]
270. r2_hidden(sK1,sK0) <- (5) [avatar component clause 268] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
271. ~1 | 5 [avatar split clause 168,268,250]
273. 6 <=> ! [X3,X4] : (r2_hidden(k6_subset_1(X3,X4),sK0) | ~r2_hidden(X3,sK0) | ~r2_hidden(X4,sK0)) [avatar definition]
274. r2_hidden(k6_subset_1(X3,X4),sK0) | ~r2_hidden(X3,sK0) | ~r2_hidden(X4,sK0) <- (6) [avatar component clause 273] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
275. 1 | 6 [avatar split clause 167,273,250]
277. 7 <=> ! [X3,X4] : (r2_hidden(k2_xboole_0(X3,X4),sK0) | ~r2_hidden(X3,sK0) | ~r2_hidden(X4,sK0)) [avatar definition]
278. r2_hidden(k2_xboole_0(X3,X4),sK0) | ~r2_hidden(X3,sK0) | ~r2_hidden(X4,sK0) <- (7) [avatar component clause 277] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
279. 1 | 7 [avatar split clause 166,277,250]
585. v3_finsub_1(sK0) | ~r2_hidden(sK9(sK0),sK0) | ~r2_hidden(sK10(sK0),sK0) <- (6) [resolution 214,274] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
586. v3_finsub_1(sK0) | ~r2_hidden(sK10(sK0),sK0) <- (6) [subsumption resolution 585,212] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
587. v3_finsub_1(sK0) <- (6) [subsumption resolution 586,213] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
589. 31 <=> v3_finsub_1(sK0) [avatar definition]
591. v3_finsub_1(sK0) <- (31) [avatar component clause 589] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
592. 31 | ~6 [avatar split clause 587,273,589]
593. v3_finsub_1(sK0) <- (1) [resolution 251,197] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
595. 31 | ~1 [avatar split clause 593,250,589]
597. 32 <=> v1_finsub_1(sK0) [avatar definition]
598. ~v1_finsub_1(sK0) <- (~32) [avatar component clause 597] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
599. v1_finsub_1(sK0) <- (32) [avatar component clause 597] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
601. v4_finsub_1(sK0) | ~v1_finsub_1(sK0) <- (31) [resolution 591,210] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
644. 38 <=> r2_hidden(sK12(sK0),sK0) [avatar definition]
646. ~r2_hidden(sK12(sK0),sK0) <- (~38) [avatar component clause 644] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
648. 39 <=> r2_hidden(sK11(sK0),sK0) [avatar definition]
650. ~r2_hidden(sK11(sK0),sK0) <- (~39) [avatar component clause 648] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
677. v1_finsub_1(sK0) <- (~38) [resolution 646,217] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
681. ~32 | 1 | ~31 [avatar split clause 601,589,250,597]
684. 32 | 38 [avatar split clause 677,644,597]
686. v1_finsub_1(sK0) <- (1) [resolution 251,196] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
687. 32 | ~1 [avatar split clause 686,250,597]
739. ~r2_hidden(sK2,sK0) | ~r2_hidden(sK1,sK0) | ~v1_finsub_1(sK0) <- (~2) [resolution 215,256] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
749. ~r2_hidden(sK1,sK0) | ~v1_finsub_1(sK0) <- (~2, 4) [subsumption resolution 739,265] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
750. ~v1_finsub_1(sK0) <- (~2, 4, 5) [subsumption resolution 749,270] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
751. $false <- (~2, 4, 5, 32) [subsumption resolution 750,599] {a:1,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
752. 2 | ~4 | ~5 | ~32 [avatar contradiction clause 751]
756. ~r2_hidden(sK2,sK0) | ~r2_hidden(sK1,sK0) | ~v3_finsub_1(sK0) <- (~3) [resolution 260,211] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
758. ~r2_hidden(sK1,sK0) | ~v3_finsub_1(sK0) <- (~3, 4) [subsumption resolution 756,265] {a:1,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
759. ~v3_finsub_1(sK0) <- (~3, 4, 5) [subsumption resolution 758,270] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
760. $false <- (~3, 4, 5, 31) [subsumption resolution 759,591] {a:1,w:0,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
761. 3 | ~4 | ~5 | ~31 [avatar contradiction clause 760]
770. ~r2_hidden(sK11(sK0),sK0) | ~r2_hidden(sK12(sK0),sK0) | v1_finsub_1(sK0) <- (7) [resolution 278,218] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
779. ~r2_hidden(sK11(sK0),sK0) | ~r2_hidden(sK12(sK0),sK0) <- (7, ~32) [subsumption resolution 770,598] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
780. ~38 | ~39 | ~7 | 32 [avatar split clause 779,597,277,648,644]
795. v1_finsub_1(sK0) <- (~39) [resolution 650,216] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
797. $false <- (~32, ~39) [subsumption resolution 795,598] {a:2,w:0,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
798. 32 | 39 [avatar contradiction clause 797]
799. $false [avatar sat refutation 261,266,271,275,279,592,595,681,684,687,752,761,780,798] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:424,thDist:-424}
% SZS output end Proof for finsub_1__t1_finsub_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 63
% >>> Preprocessing
% Introduced skolems: 20
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 83
% Generated clauses: 470
% Active clauses: 134
% Passive clauses: 254
% Final active clauses: 109
% Final passive clauses: 112
% Main loop iterations started: 138
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 31
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 32
% Forward subsumptions: 78
% 
% >>> Generating Inferences
% Binary resolution: 209
% Unit resulting resolution: 7
% Forward superposition: 34
% Backward superposition: 22
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 25
% Sat splitting refutations: 14
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 67
% SAT solver unit clauses: 10
% SAT solver binary clauses: 38
% 
% Memory used [KB]: 5373
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% forward subsumption: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

