i: [1,63,0,2,1,1,0,0] t119_funct_4
a: [3,100,0,2,1,63]
s: 100
i: [1,62,0,2,1,1,0,0] t119_funct_4
a: [3,105,0,2,1,62]
s: 105
i: [1,61,0,2,1,1,0,0] t119_funct_4
a: [3,110,0,2,1,61]
s: 110
i: [1,60,0,2,1,1,0,0] t119_funct_4
a: [3,115,0,2,1,60]
s: 115
i: [1,64,0,5,1,1,0,0] t119_funct_4
a: [3,95,0,5,1,64]
s: 95
i: [1,65,0,7,1,1,0,0] t119_funct_4
a: [3,90,0,7,1,65]
s: 90
i: [1,86,0,2,1,0,0,1] rc1_funct_1
a: [3,127,0,2,1,86]
s: 127
i: [1,85,0,2,1,0,0,1] rc1_funct_1
a: [3,132,0,2,1,85]
s: 132
i: [1,73,0,3,1,0,0,255] reflexivity_r1_tarski
s: 73
i: [1,72,0,4,1,0,0,2] existence_m1_subset_1
s: 72
i: [1,74,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 74
i: [1,75,0,6,1,0,0,1] dt_k6_subset_1
s: 75
i: [1,78,0,6,2,0,0,2] fc2_relat_1
s: 78
i: [1,79,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 79
i: [1,71,0,7,2,0,0,255] t3_subset
s: 71
i: [1,76,0,7,1,0,0,1] redefinition_k6_subset_1
s: 76
r: 78
d: [2,134,1,5,1,1,40,79,95]
a: [3,138,1,5,1,134]
s: 138
d: [2,140,0,6,2,0,23,78,76]
s: 140
i: [1,77,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 77
i: [1,67,0,8,3,0,0,2] cc2_relat_1
s: 67
i: [1,84,0,8,3,0,0,2] fc3_relat_1
s: 84
d: [2,142,1,6,2,0,40,67,72]
s: 142
i: [1,80,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,122,0,9,3,80]
s: 122
d: [2,144,1,7,3,0,40,67,71]
s: 144
i: [1,68,0,10,4,0,0,2] cc3_funct_1
s: 68
i: [1,70,0,10,2,0,0,2] t88_xboole_1
s: 70
d: [2,150,1,8,3,0,40,68,72]
s: 150
d: [2,153,0,10,2,0,22,70,76]
s: 153
d: [2,151,1,8,3,0,40,68,75]
s: 151
i: [1,66,0,12,4,0,0,1] t179_relat_1
s: 66
i: [1,82,0,12,5,0,0,1] dt_k1_funct_4
s: 82
d: [2,160,1,7,3,1,40,66,95]
d: [2,162,1,5,2,2,21,160,105]
d: [2,163,1,3,1,3,21,162,115]
a: [3,167,1,3,1,163]
s: 167
i: [1,83,0,12,5,0,0,1] dt_k1_funct_4
s: 83
d: [2,161,2,7,3,1,40,66,138]
d: [2,169,2,5,2,2,21,161,115]
d: [2,170,2,3,1,3,21,169,105]
a: [3,174,2,3,1,170]
s: 174
i: [1,81,0,17,5,0,0,1] rd3_funct_4
s: 81
i: [1,69,0,20,6,0,0,1] t31_funct_4
s: 69
d: [2,154,1,8,2,0,44,153,74]
s: 154
d: [2,152,1,9,4,0,40,68,71]
s: 152
d: [2,205,2,7,2,0,44,75,154]
s: 205
d: [2,157,1,9,3,0,44,140,153]
s: 157
d: [2,158,1,9,2,0,44,75,153]
s: 158
d: [2,212,2,9,3,0,44,157,77]
s: 212
d: [2,155,1,10,2,0,44,153,77]
s: 155
d: [2,213,2,9,3,0,44,157,77]
s: 213
d: [2,156,1,10,2,0,44,153,77]
s: 156
d: [2,192,1,20,6,0,44,69,77]
s: 192
d: [2,218,2,9,2,0,44,158,77]
s: 218
d: [2,193,1,20,6,0,44,69,77]
s: 193
d: [2,219,2,9,2,0,44,158,77]
s: 219
d: [2,197,1,20,6,0,44,77,69]
s: 197
d: [2,198,1,20,6,0,44,77,69]
s: 198
d: [2,232,2,9,2,0,44,75,155]
s: 232
d: [2,195,1,23,7,0,44,153,69]
d: [2,202,1,20,6,0,21,195,66]
s: 202
d: [2,263,2,0,0,6,52,115,110,105,100,95,90,202]
e: 263
d: [2,266,2,16,6,1,44,90,202]
d: [2,270,2,13,5,1,19,266]
d: [2,276,2,8,4,2,21,270,95]
d: [2,277,2,6,3,3,21,276,100]
d: [2,278,2,4,2,4,21,277,105]
d: [2,279,2,2,1,5,21,278,110]
d: [2,280,2,0,0,6,21,279,115]
e: 280
f: 280,60,61,62,63,64,65
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t119_funct_4
% SZS output start Proof for funct_4__t119_funct_4
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k6_subset_1(k1_funct_4(X0,X1),X1) = X0))) [input t119_funct_4]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k6_subset_1(k1_funct_4(X0,X1),X1) = X0))) [negated conjecture 1]
20. ! [X0,X1] : k6_subset_1(X0,X1) = k4_xboole_0(X0,X1) [input redefinition_k6_subset_1]
23. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => r1_xboole_0(X0,X1)))) [input t179_relat_1]
24. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k1_funct_4(X0,X1) = k2_xboole_0(X0,X1)))) [input t31_funct_4]
26. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) => k4_xboole_0(k2_xboole_0(X0,X1),X1) = X0) [input t88_xboole_1]
27. ! [X0,X1] : (r1_xboole_0(X0,X1) => k4_xboole_0(k2_xboole_0(X0,X1),X1) = X0) [flattening 26]
32. ? [X0] : (? [X1] : ((k6_subset_1(k1_funct_4(X0,X1),X1) != X0 & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) & (v1_funct_1(X1) & v1_relat_1(X1))) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
33. ? [X0] : (? [X1] : (k6_subset_1(k1_funct_4(X0,X1),X1) != X0 & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 32]
34. ! [X0] : (! [X1] : ((r1_xboole_0(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 23]
35. ! [X0] : (! [X1] : (r1_xboole_0(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 34]
39. ! [X0] : (! [X1] : ((k1_funct_4(X0,X1) = k2_xboole_0(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 24]
40. ! [X0] : (! [X1] : (k1_funct_4(X0,X1) = k2_xboole_0(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 39]
41. ! [X0,X1] : (k4_xboole_0(k2_xboole_0(X0,X1),X1) = X0 | ~r1_xboole_0(X0,X1)) [ennf transformation 27]
53. ? [X0] : (? [X1] : (k6_subset_1(k1_funct_4(X0,X1),X1) != X0 & r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : (sK0 != k6_subset_1(k1_funct_4(sK0,X1),X1) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
54. ? [X1] : (sK0 != k6_subset_1(k1_funct_4(sK0,X1),X1) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(X1)) & v1_funct_1(X1) & v1_relat_1(X1)) => (sK0 != k6_subset_1(k1_funct_4(sK0,sK1),sK1) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) [choice axiom]
55. (sK0 != k6_subset_1(k1_funct_4(sK0,sK1),sK1) & r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) & v1_funct_1(sK1) & v1_relat_1(sK1)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 33,54,53]
60. v1_relat_1(sK0) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
61. v1_funct_1(sK0) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
62. v1_relat_1(sK1) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
63. v1_funct_1(sK1) [cnf transformation 55] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
64. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) [cnf transformation 55] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. sK0 != k6_subset_1(k1_funct_4(sK0,sK1),sK1) [cnf transformation 55] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
66. ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | r1_xboole_0(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. k1_funct_4(X0,X1) = k2_xboole_0(X0,X1) | ~r1_xboole_0(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 40] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. k4_xboole_0(k2_xboole_0(X0,X1),X1) = X0 | ~r1_xboole_0(X0,X1) [cnf transformation 41] {a:0,w:10,wCS:4000,sine:2,thAx:0,allAx:1,thDist:-1}
76. k6_subset_1(X0,X1) = k4_xboole_0(X0,X1) [cnf transformation 20] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
88. 1 <=> sK0 = k6_subset_1(k1_funct_4(sK0,sK1),sK1) [avatar definition]
90. sK0 != k6_subset_1(k1_funct_4(sK0,sK1),sK1) <- (~1) [avatar component clause 88] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
91. ~1 [avatar split clause 65,88]
93. 2 <=> r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) [avatar definition]
95. r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) <- (2) [avatar component clause 93] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
96. 2 [avatar split clause 64,93]
98. 3 <=> v1_funct_1(sK1) [avatar definition]
100. v1_funct_1(sK1) <- (3) [avatar component clause 98] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
101. 3 [avatar split clause 63,98]
103. 4 <=> v1_relat_1(sK1) [avatar definition]
105. v1_relat_1(sK1) <- (4) [avatar component clause 103] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
106. 4 [avatar split clause 62,103]
108. 5 <=> v1_funct_1(sK0) [avatar definition]
110. v1_funct_1(sK0) <- (5) [avatar component clause 108] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
111. 5 [avatar split clause 61,108]
113. 6 <=> v1_relat_1(sK0) [avatar definition]
115. v1_relat_1(sK0) <- (6) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. 6 [avatar split clause 60,113]
153. k6_subset_1(k2_xboole_0(X0,X1),X1) = X0 | ~r1_xboole_0(X0,X1) [forward demodulation 70,76] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
195. k6_subset_1(k1_funct_4(X1,X2),X2) = X1 | ~r1_xboole_0(X1,X2) | ~r1_xboole_0(k9_xtuple_0(X1),k9_xtuple_0(X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [superposition 153,69] {a:1,w:23,wCS:9200,sine:1,thAx:0,allAx:2,thDist:-2}
202. k6_subset_1(k1_funct_4(X1,X2),X2) = X1 | ~r1_xboole_0(k9_xtuple_0(X1),k9_xtuple_0(X2)) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_funct_1(X1) | ~v1_relat_1(X1) [subsumption resolution 195,66] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
266. sK0 != sK0 | ~r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [superposition 90,202] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
270. ~r1_xboole_0(k9_xtuple_0(sK0),k9_xtuple_0(sK1)) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1) [trivial inequality removal 266] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
276. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 270,95] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
277. ~v1_relat_1(sK1) | ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3) [subsumption resolution 276,100] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
278. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) <- (~1, 2, 3, 4) [subsumption resolution 277,105] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
279. ~v1_relat_1(sK0) <- (~1, 2, 3, 4, 5) [subsumption resolution 278,110] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
280. $false <- (~1, 2, 3, 4, 5, 6) [subsumption resolution 279,115] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
281. 1 | ~2 | ~3 | ~4 | ~5 | ~6 [avatar contradiction clause 280]
282. $false [avatar sat refutation 91,96,101,106,111,116,281] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:88,thDist:-88}
% SZS output end Proof for funct_4__t119_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% >>> Preprocessing
% Introduced skolems: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 169
% Active clauses: 45
% Passive clauses: 66
% Final active clauses: 44
% Final passive clauses: 13
% Main loop iterations started: 53
% 
% >>> Simplifying Inferences
% Duplicate literals: 45
% Trivial inequalities: 1
% Fw subsumption resolutions: 17
% Fw demodulations: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 21
% Forward subsumptions: 36
% 
% >>> Generating Inferences
% Binary resolution: 22
% Unit resulting resolution: 1
% Forward superposition: 40
% Backward superposition: 29
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 8
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

