i: [1,109,0,2,1,1,0,0] t10_diraf
a: [3,173,0,2,1,109]
s: 173
i: [1,108,0,2,1,1,0,0] t10_diraf
a: [3,178,0,2,1,108]
s: 178
i: [1,107,0,2,1,1,0,0] t10_diraf
a: [3,183,0,2,1,107]
s: 183
i: [1,111,0,4,1,1,0,0] t10_diraf
a: [3,163,0,4,1,111]
s: 163
i: [1,110,0,4,1,1,0,0] t10_diraf
a: [3,168,0,4,1,110]
s: 168
i: [1,112,0,10,2,1,0,0] t10_diraf
a: [3,154,0,5,1,112]
s: 154
i: [1,113,0,2,1,0,0,255] spc1_boole
a: [3,188,0,2,1,113]
s: 188
i: [1,114,0,2,1,0,0,255] fc1_xboole_0
a: [3,193,0,2,1,114]
s: 193
i: [1,147,0,2,1,0,0,2] rc2_xboole_0
a: [3,198,0,2,1,147]
s: 198
i: [1,148,0,2,1,0,0,1] existence_l1_analoaf
a: [3,203,0,2,1,148]
s: 203
i: [1,149,0,2,1,0,0,2] existence_l1_struct_0
a: [3,208,0,2,1,149]
s: 208
i: [1,150,0,2,1,0,0,2] rc1_xboole_0
a: [3,213,0,2,1,150]
s: 213
i: [1,115,0,4,2,0,0,1] dt_l1_analoaf
s: 115
d: [2,215,1,2,1,1,40,115,173]
a: [3,220,1,2,1,215]
s: 220
i: [1,145,0,4,1,0,0,1] existence_m1_subset_1
s: 145
i: [1,130,0,5,2,0,0,255] t6_boole
s: 130
d: [2,216,1,2,1,1,40,115,203]
a: [3,225,1,2,1,216]
s: 225
i: [1,144,0,5,2,0,0,255] t7_boole
s: 144
i: [1,116,0,6,3,0,0,255] cc7_struct_0
s: 116
i: [1,117,0,6,3,0,0,255] cc5_struct_0
s: 117
i: [1,119,0,6,3,0,0,1] cc17_struct_0
s: 119
d: [2,229,1,4,2,1,40,116,208]
a: [3,235,1,2,1,229]
s: 235
i: [1,120,0,6,3,0,0,1] cc2_struct_0
s: 120
d: [2,231,2,4,2,1,40,116,225]
a: [3,244,2,2,1,231]
s: 244
i: [1,121,0,6,3,0,0,1] cc1_struct_0
s: 121
i: [1,124,0,6,3,0,0,255] cc4_struct_0
s: 124
i: [1,125,0,6,3,0,0,255] cc6_struct_0
s: 125
i: [1,141,0,6,2,0,0,255] t1_subset
s: 141
i: [1,146,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 146
i: [1,122,0,7,3,0,0,255] cc8_struct_0
s: 122
i: [1,126,0,7,3,0,0,255] cc11_struct_0
s: 126
i: [1,127,0,7,3,0,0,255] cc11_struct_0
s: 127
i: [1,128,0,7,3,0,0,255] cc9_struct_0
s: 128
i: [1,131,0,7,3,0,0,255] fc9_struct_0
s: 131
i: [1,134,0,7,3,0,0,1] fc2_struct_0
s: 134
i: [1,137,0,7,3,0,0,255] fc6_struct_0
s: 137
i: [1,138,0,7,3,0,0,255] fc8_struct_0
s: 138
i: [1,139,0,7,3,0,0,1] fc1_struct_0
s: 139
i: [1,140,0,7,3,0,0,255] fc7_struct_0
s: 140
i: [1,143,0,7,3,0,0,255] t8_boole
s: 143
i: [1,142,0,8,3,0,0,255] t2_subset
s: 142
d: [2,260,1,7,2,1,40,142,163]
a: [3,270,1,3,1,260]
s: 270
r: 270
i: [1,129,0,9,4,0,0,255] cc10_struct_0
s: 129
a: [3,269,1,3,1,260]
s: 269
i: [1,135,0,24,7,0,0,1] t4_diraf
s: 135
i: [1,136,0,24,7,0,0,1] t4_diraf
s: 136
d: [2,277,2,4,2,1,40,270,134]
d: [2,278,2,2,1,2,21,277,220]
a: [3,281,2,2,1,278]
s: 281
i: [1,132,0,27,7,0,0,1] d3_diraf
s: 132
a: [3,266,1,4,1,260]
s: 266
i: [1,133,0,27,7,0,0,1] d3_diraf
s: 133
d: [2,261,1,7,2,1,40,142,168]
a: [3,275,1,4,1,261]
s: 275
d: [2,291,2,4,1,1,40,266,146]
a: [3,297,2,4,1,291]
s: 297
d: [2,262,1,6,2,0,40,142,145]
s: 262
d: [2,305,2,4,1,1,40,275,146]
a: [3,311,2,4,1,305]
s: 311
d: [2,227,1,7,3,0,44,130,130]
s: 227
d: [2,250,1,9,4,0,44,122,130]
s: 250
d: [2,313,2,6,2,0,40,262,146]
s: 313
d: [2,252,1,9,4,0,44,128,130]
s: 252
d: [2,300,1,39,12,0,40,133,136]
d: [2,303,1,21,7,0,18,300]
s: 303
d: [2,299,1,39,12,0,40,133,135]
d: [2,304,1,21,7,0,18,299]
s: 304
d: [2,328,3,0,0,7,52,183,178,281,173,168,163,154,304]
e: 328
d: [2,329,2,16,6,1,40,304,154]
d: [2,331,2,12,5,2,21,329,163]
d: [2,332,2,8,4,3,21,331,168]
d: [2,333,2,6,3,4,21,332,173]
d: [2,334,2,4,2,5,21,333,281]
d: [2,335,2,2,1,6,21,334,178]
d: [2,336,2,0,0,7,21,335,183]
e: 336
r: 154
a: [3,158,0,5,1,112]
s: 158
d: [2,338,3,0,0,7,52,183,178,281,173,168,163,158,303]
e: 338
d: [2,339,2,16,6,1,40,158,303]
d: [2,341,2,12,5,2,21,339,163]
d: [2,342,2,8,4,3,21,341,168]
d: [2,343,2,6,3,4,21,342,173]
d: [2,344,2,4,2,5,21,343,281]
d: [2,345,2,2,1,6,21,344,178]
d: [2,346,2,0,0,7,21,345,183]
e: 346
r: 269
r: 281
r: 266
r: 275
a: [3,282,2,2,1,278]
s: 282
d: [2,353,3,0,0,3,52,220,183,282,119]
e: 353
d: [2,354,3,4,2,1,40,282,119]
d: [2,357,3,2,1,2,21,354,183]
d: [2,358,3,0,0,3,21,357,220]
e: 358
d: [2,350,2,4,2,5,21,343,178]
d: [2,351,2,2,1,6,21,350,183]
f: 358,351,336,215,107,108,109,110,111,112
% Refutation found. Thanks to Tanya!
% SZS status Theorem for diraf__t10_diraf
% SZS output start Proof for diraf__t10_diraf
1. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r1_diraf(X0,X1,X2,X2) & r1_diraf(X0,X1,X1,X2))))) [input t10_diraf]
2. ~! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r1_diraf(X0,X1,X2,X2) & r1_diraf(X0,X1,X1,X2))))) [negated conjecture 1]
6. ! [X0] : (l1_struct_0(X0) => (~v7_struct_0(X0) => ~v2_struct_0(X0))) [input cc17_struct_0]
15. ! [X0] : ((l1_analoaf(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_diraf(X0,X1,X2,X3) <=> r2_analoaf(X0,X1,X2,X2,X3)))))) [input d3_diraf]
17. ! [X0] : (l1_analoaf(X0) => l1_struct_0(X0)) [input dt_l1_analoaf]
36. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r2_analoaf(X0,X2,X3,X1,X1) & r2_analoaf(X0,X1,X1,X2,X3)))))) [input t4_diraf]
44. ? [X0] : (? [X1] : (? [X2] : ((~r1_diraf(X0,X1,X2,X2) | ~r1_diraf(X0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
45. ? [X0] : (? [X1] : (? [X2] : ((~r1_diraf(X0,X1,X2,X2) | ~r1_diraf(X0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) [flattening 44]
46. ! [X0] : (l1_struct_0(X0) | ~l1_analoaf(X0)) [ennf transformation 17]
51. ! [X0] : ((~v2_struct_0(X0) | v7_struct_0(X0)) | ~l1_struct_0(X0)) [ennf transformation 6]
52. ! [X0] : (~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0)) [flattening 51]
72. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_diraf(X0,X1,X2,X3) <=> r2_analoaf(X0,X1,X2,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | v2_struct_0(X0))) [ennf transformation 15]
73. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_diraf(X0,X1,X2,X3) <=> r2_analoaf(X0,X1,X2,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 72]
76. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_analoaf(X0,X2,X3,X1,X1) & r2_analoaf(X0,X1,X1,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0))) [ennf transformation 36]
77. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_analoaf(X0,X2,X3,X1,X1) & r2_analoaf(X0,X1,X1,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0)) [flattening 76]
92. ? [X0] : (? [X1] : (? [X2] : ((~r1_diraf(X0,X1,X2,X2) | ~r1_diraf(X0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : ((~r1_diraf(sK0,X1,X2,X2) | ~r1_diraf(sK0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
93. ? [X1] : (? [X2] : ((~r1_diraf(sK0,X1,X2,X2) | ~r1_diraf(sK0,X1,X1,X2)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : ((~r1_diraf(sK0,sK1,X2,X2) | ~r1_diraf(sK0,sK1,sK1,X2)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
94. ? [X2] : ((~r1_diraf(sK0,sK1,X2,X2) | ~r1_diraf(sK0,sK1,sK1,X2)) & m1_subset_1(X2,u1_struct_0(sK0))) => ((~r1_diraf(sK0,sK1,sK2,sK2) | ~r1_diraf(sK0,sK1,sK1,sK2)) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
95. (((~r1_diraf(sK0,sK1,sK2,sK2) | ~r1_diraf(sK0,sK1,sK1,sK2)) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0) [skolemisation 45,94,93,92]
96. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r1_diraf(X0,X1,X2,X3) | ~r2_analoaf(X0,X1,X2,X2,X3)) & (r2_analoaf(X0,X1,X2,X2,X3) | ~r1_diraf(X0,X1,X2,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [nnf transformation 73]
107. ~v7_struct_0(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
108. v2_analoaf(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
109. l1_analoaf(sK0) [cnf transformation 95] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
110. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 95] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
111. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 95] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
112. ~r1_diraf(sK0,sK1,sK2,sK2) | ~r1_diraf(sK0,sK1,sK1,sK2) [cnf transformation 95] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
115. ~l1_analoaf(X0) | l1_struct_0(X0) [cnf transformation 46] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
119. ~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 52] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
133. ~r2_analoaf(X0,X1,X2,X2,X3) | r1_diraf(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 96] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
135. r2_analoaf(X0,X1,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 77] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
136. r2_analoaf(X0,X2,X3,X1,X1) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 77] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
152. 1 <=> r1_diraf(sK0,sK1,sK1,sK2) [avatar definition]
154. ~r1_diraf(sK0,sK1,sK1,sK2) <- (~1) [avatar component clause 152] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
156. 2 <=> r1_diraf(sK0,sK1,sK2,sK2) [avatar definition]
158. ~r1_diraf(sK0,sK1,sK2,sK2) <- (~2) [avatar component clause 156] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
159. ~1 | ~2 [avatar split clause 112,156,152]
161. 3 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
163. m1_subset_1(sK2,u1_struct_0(sK0)) <- (3) [avatar component clause 161] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
164. 3 [avatar split clause 111,161]
166. 4 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
168. m1_subset_1(sK1,u1_struct_0(sK0)) <- (4) [avatar component clause 166] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
169. 4 [avatar split clause 110,166]
171. 5 <=> l1_analoaf(sK0) [avatar definition]
173. l1_analoaf(sK0) <- (5) [avatar component clause 171] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
174. 5 [avatar split clause 109,171]
176. 6 <=> v2_analoaf(sK0) [avatar definition]
178. v2_analoaf(sK0) <- (6) [avatar component clause 176] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
179. 6 [avatar split clause 108,176]
181. 7 <=> v7_struct_0(sK0) [avatar definition]
183. ~v7_struct_0(sK0) <- (~7) [avatar component clause 181] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
184. ~7 [avatar split clause 107,181]
215. l1_struct_0(sK0) <- (5) [resolution 115,173] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
218. 14 <=> l1_struct_0(sK0) [avatar definition]
220. l1_struct_0(sK0) <- (14) [avatar component clause 218] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
221. 14 | ~5 [avatar split clause 215,171,218]
280. 23 <=> v2_struct_0(sK0) [avatar definition]
281. ~v2_struct_0(sK0) <- (~23) [avatar component clause 280] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
282. v2_struct_0(sK0) <- (23) [avatar component clause 280] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
299. r1_diraf(X0,X1,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [resolution 133,135] {a:1,w:39,wCS:15600,sine:1,thAx:0,allAx:2,thDist:-2}
300. r1_diraf(X3,X4,X5,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_analoaf(X3) | v2_struct_0(X3) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~l1_analoaf(X3) | ~v2_analoaf(X3) | v7_struct_0(X3) [resolution 133,136] {a:1,w:39,wCS:15600,sine:1,thAx:0,allAx:2,thDist:-2}
303. r1_diraf(X3,X4,X5,X5) | ~m1_subset_1(X5,u1_struct_0(X3)) | ~m1_subset_1(X4,u1_struct_0(X3)) | ~l1_analoaf(X3) | v2_struct_0(X3) | ~v2_analoaf(X3) | v7_struct_0(X3) [duplicate literal removal 300] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
304. r1_diraf(X0,X1,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [duplicate literal removal 299] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
329. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1) [resolution 304,154] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
331. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 3) [subsumption resolution 329,163] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
332. ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 3, 4) [subsumption resolution 331,168] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
333. v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 3, 4, 5) [subsumption resolution 332,173] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
334. ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~1, 3, 4, 5, ~23) [subsumption resolution 333,281] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
335. v7_struct_0(sK0) <- (~1, 3, 4, 5, 6, ~23) [subsumption resolution 334,178] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
336. $false <- (~1, 3, 4, 5, 6, ~7, ~23) [subsumption resolution 335,183] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
337. 1 | ~3 | ~4 | ~5 | ~6 | 7 | 23 [avatar contradiction clause 336]
339. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2) [resolution 158,303] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
341. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 3) [subsumption resolution 339,163] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
342. ~l1_analoaf(sK0) | v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 3, 4) [subsumption resolution 341,168] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
343. v2_struct_0(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (~2, 3, 4, 5) [subsumption resolution 342,173] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
350. v2_struct_0(sK0) | v7_struct_0(sK0) <- (~2, 3, 4, 5, 6) [subsumption resolution 343,178] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
351. v2_struct_0(sK0) <- (~2, 3, 4, 5, 6, ~7) [subsumption resolution 350,183] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
352. 23 | 2 | ~3 | ~4 | ~5 | ~6 | 7 [avatar split clause 351,181,176,171,166,161,156,280]
354. v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (23) [resolution 282,119] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
357. ~l1_struct_0(sK0) <- (~7, 23) [subsumption resolution 354,183] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
358. $false <- (~7, 14, 23) [subsumption resolution 357,220] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
359. 7 | ~14 | ~23 [avatar contradiction clause 358]
361. $false [avatar sat refutation 159,164,169,174,179,184,221,337,352,359] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:250,thDist:-250}
% SZS output end Proof for diraf__t10_diraf
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 44
% Generated clauses: 154
% Active clauses: 57
% Passive clauses: 65
% Final active clauses: 51
% Final passive clauses: 3
% Main loop iterations started: 62
% 
% >>> Simplifying Inferences
% Duplicate literals: 30
% Fw subsumption resolutions: 19
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 14
% 
% >>> Generating Inferences
% Binary resolution: 38
% Unit resulting resolution: 3
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 33
% SAT solver unit clauses: 11
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

