i: [1,30,0,2,1,1,0,0] t54_wellord1
a: [3,49,0,2,1,30]
s: 49
i: [1,29,0,2,1,1,0,0] t54_wellord1
a: [3,54,0,2,1,29]
s: 54
i: [1,27,0,2,1,1,0,0] t54_wellord1
a: [3,64,0,2,1,27]
s: 64
i: [1,26,0,2,1,1,0,0] t54_wellord1
a: [3,69,0,2,1,26]
s: 69
i: [1,28,0,3,1,1,0,0] t54_wellord1
a: [3,59,0,3,1,28]
s: 59
i: [1,45,0,2,1,0,0,2] rc1_funct_1
a: [3,74,0,2,1,45]
s: 74
i: [1,44,0,2,1,0,0,2] rc1_funct_1
a: [3,79,0,2,1,44]
s: 79
i: [1,39,0,6,3,0,0,255] cc1_wellord1
s: 39
d: [2,81,1,4,2,1,40,39,54]
d: [2,82,1,2,1,2,21,81,69]
a: [3,86,1,2,1,82]
s: 86
i: [1,40,0,6,3,0,0,255] cc1_wellord1
s: 40
i: [1,41,0,6,3,0,0,255] cc1_wellord1
s: 41
d: [2,88,1,4,2,1,40,40,54]
d: [2,89,1,2,1,2,21,88,69]
a: [3,93,1,2,1,89]
s: 93
i: [1,42,0,6,3,0,0,255] cc1_wellord1
s: 42
d: [2,95,1,4,2,1,40,41,54]
d: [2,96,1,2,1,2,21,95,69]
a: [3,100,1,2,1,96]
s: 100
i: [1,43,0,6,3,0,0,255] cc1_wellord1
s: 43
i: [1,31,0,11,4,0,0,1] d8_wellord1
s: 31
d: [2,102,1,4,2,1,40,42,54]
d: [2,103,1,2,1,2,21,102,69]
a: [3,107,1,2,1,103]
s: 107
i: [1,32,0,11,4,0,0,1] d8_wellord1
s: 32
d: [2,109,1,4,2,1,40,43,54]
d: [2,110,1,2,1,2,21,109,69]
a: [3,114,1,2,1,110]
s: 114
i: [1,33,0,13,4,0,0,1] d8_wellord1
s: 33
i: [1,37,0,14,7,0,0,255] cc2_wellord1
s: 37
i: [1,34,0,15,6,0,0,1] d8_wellord1
s: 34
i: [1,35,0,16,7,0,0,2] t44_wellord1
s: 35
d: [2,119,1,23,9,0,40,35,33]
d: [2,120,1,19,7,0,18,119]
d: [2,121,1,15,6,0,21,120,32]
d: [2,122,1,11,5,0,21,121,31]
s: 122
d: [2,123,2,0,0,5,52,49,69,64,54,59,122]
e: 123
d: [2,124,2,8,4,1,40,122,59]
d: [2,126,2,6,3,2,21,124,54]
d: [2,127,2,4,2,3,21,126,64]
d: [2,128,2,2,1,4,21,127,69]
d: [2,129,2,0,0,5,21,128,49]
e: 129
f: 129,26,27,28,29,30
% Refutation found. Thanks to Tanya!
% SZS status Theorem for wellord1__t54_wellord1
% SZS output start Proof for wellord1__t54_wellord1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ((v2_wellord1(X0) & r4_wellord1(X0,X1)) => v2_wellord1(X1)))) [input t54_wellord1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ((v2_wellord1(X0) & r4_wellord1(X0,X1)) => v2_wellord1(X1)))) [negated conjecture 1]
5. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r4_wellord1(X0,X1) <=> ? [X2] : (r3_wellord1(X0,X1,X2) & v1_funct_1(X2) & v1_relat_1(X2))))) [input d8_wellord1]
7. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2] : ((v1_funct_1(X2) & v1_relat_1(X2)) => ((r3_wellord1(X0,X1,X2) & v2_wellord1(X0)) => v2_wellord1(X1))))) [input t44_wellord1]
8. ? [X0] : (? [X1] : ((~v2_wellord1(X1) & (v2_wellord1(X0) & r4_wellord1(X0,X1))) & v1_relat_1(X1)) & v1_relat_1(X0)) [ennf transformation 2]
9. ? [X0] : (? [X1] : (~v2_wellord1(X1) & v2_wellord1(X0) & r4_wellord1(X0,X1) & v1_relat_1(X1)) & v1_relat_1(X0)) [flattening 8]
10. ! [X0] : (! [X1] : ((r4_wellord1(X0,X1) <=> ? [X2] : (r3_wellord1(X0,X1,X2) & v1_funct_1(X2) & v1_relat_1(X2))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 5]
11. ! [X0] : (! [X1] : (! [X2] : ((v2_wellord1(X1) | (~r3_wellord1(X0,X1,X2) | ~v2_wellord1(X0))) | (~v1_funct_1(X2) | ~v1_relat_1(X2))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 7]
12. ! [X0] : (! [X1] : (! [X2] : (v2_wellord1(X1) | ~r3_wellord1(X0,X1,X2) | ~v2_wellord1(X0) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 11]
17. ? [X0] : (? [X1] : (~v2_wellord1(X1) & v2_wellord1(X0) & r4_wellord1(X0,X1) & v1_relat_1(X1)) & v1_relat_1(X0)) => (? [X1] : (~v2_wellord1(X1) & v2_wellord1(sK0) & r4_wellord1(sK0,X1) & v1_relat_1(X1)) & v1_relat_1(sK0)) [choice axiom]
18. ? [X1] : (~v2_wellord1(X1) & v2_wellord1(sK0) & r4_wellord1(sK0,X1) & v1_relat_1(X1)) => (~v2_wellord1(sK1) & v2_wellord1(sK0) & r4_wellord1(sK0,sK1) & v1_relat_1(sK1)) [choice axiom]
19. (~v2_wellord1(sK1) & v2_wellord1(sK0) & r4_wellord1(sK0,sK1) & v1_relat_1(sK1)) & v1_relat_1(sK0) [skolemisation 9,18,17]
20. ! [X0] : (! [X1] : (((r4_wellord1(X0,X1) | ! [X2] : (~r3_wellord1(X0,X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X2] : (r3_wellord1(X0,X1,X2) & v1_funct_1(X2) & v1_relat_1(X2)) | ~r4_wellord1(X0,X1))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [nnf transformation 10]
21. ! [X0] : (! [X1] : (((r4_wellord1(X0,X1) | ! [X2] : (~r3_wellord1(X0,X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X3] : (r3_wellord1(X0,X1,X3) & v1_funct_1(X3) & v1_relat_1(X3)) | ~r4_wellord1(X0,X1))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [rectify 20]
22. ! [X1,X0] : (? [X3] : (r3_wellord1(X0,X1,X3) & v1_funct_1(X3) & v1_relat_1(X3)) => (r3_wellord1(X0,X1,sK2(X0,X1)) & v1_funct_1(sK2(X0,X1)) & v1_relat_1(sK2(X0,X1)))) [choice axiom]
23. ! [X0] : (! [X1] : (((r4_wellord1(X0,X1) | ! [X2] : (~r3_wellord1(X0,X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & ((r3_wellord1(X0,X1,sK2(X0,X1)) & v1_funct_1(sK2(X0,X1)) & v1_relat_1(sK2(X0,X1))) | ~r4_wellord1(X0,X1))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [skolemisation 21,22]
26. v1_relat_1(sK0) [cnf transformation 19] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
27. v1_relat_1(sK1) [cnf transformation 19] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
28. r4_wellord1(sK0,sK1) [cnf transformation 19] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
29. v2_wellord1(sK0) [cnf transformation 19] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
30. ~v2_wellord1(sK1) [cnf transformation 19] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
31. v1_relat_1(sK2(X0,X1)) | ~r4_wellord1(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
32. v1_funct_1(sK2(X0,X1)) | ~r4_wellord1(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
33. r3_wellord1(X0,X1,sK2(X0,X1)) | ~r4_wellord1(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
35. ~r3_wellord1(X0,X1,X2) | v2_wellord1(X1) | ~v2_wellord1(X0) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 12] {a:0,w:16,wCS:6400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
47. 1 <=> v2_wellord1(sK1) [avatar definition]
49. ~v2_wellord1(sK1) <- (~1) [avatar component clause 47] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
50. ~1 [avatar split clause 30,47]
52. 2 <=> v2_wellord1(sK0) [avatar definition]
54. v2_wellord1(sK0) <- (2) [avatar component clause 52] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
55. 2 [avatar split clause 29,52]
57. 3 <=> r4_wellord1(sK0,sK1) [avatar definition]
59. r4_wellord1(sK0,sK1) <- (3) [avatar component clause 57] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
60. 3 [avatar split clause 28,57]
62. 4 <=> v1_relat_1(sK1) [avatar definition]
64. v1_relat_1(sK1) <- (4) [avatar component clause 62] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
65. 4 [avatar split clause 27,62]
67. 5 <=> v1_relat_1(sK0) [avatar definition]
69. v1_relat_1(sK0) <- (5) [avatar component clause 67] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
70. 5 [avatar split clause 26,67]
119. v2_wellord1(X0) | ~v2_wellord1(X1) | ~v1_funct_1(sK2(X1,X0)) | ~v1_relat_1(sK2(X1,X0)) | ~v1_relat_1(X0) | ~v1_relat_1(X1) | ~r4_wellord1(X1,X0) | ~v1_relat_1(X0) | ~v1_relat_1(X1) [resolution 35,33] {a:1,w:23,wCS:9200,sine:1,thAx:0,allAx:3,thDist:-3}
120. v2_wellord1(X0) | ~v2_wellord1(X1) | ~v1_funct_1(sK2(X1,X0)) | ~v1_relat_1(sK2(X1,X0)) | ~v1_relat_1(X0) | ~v1_relat_1(X1) | ~r4_wellord1(X1,X0) [duplicate literal removal 119] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:3,thDist:-3}
121. v2_wellord1(X0) | ~v2_wellord1(X1) | ~v1_relat_1(sK2(X1,X0)) | ~v1_relat_1(X0) | ~v1_relat_1(X1) | ~r4_wellord1(X1,X0) [subsumption resolution 120,32] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:3,thDist:-3}
122. ~r4_wellord1(X1,X0) | ~v2_wellord1(X1) | ~v1_relat_1(X0) | ~v1_relat_1(X1) | v2_wellord1(X0) [subsumption resolution 121,31] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
124. ~v2_wellord1(sK0) | ~v1_relat_1(sK1) | ~v1_relat_1(sK0) | v2_wellord1(sK1) <- (3) [resolution 122,59] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
126. ~v1_relat_1(sK1) | ~v1_relat_1(sK0) | v2_wellord1(sK1) <- (2, 3) [subsumption resolution 124,54] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
127. ~v1_relat_1(sK0) | v2_wellord1(sK1) <- (2, 3, 4) [subsumption resolution 126,64] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
128. v2_wellord1(sK1) <- (2, 3, 4, 5) [subsumption resolution 127,69] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
129. $false <- (~1, 2, 3, 4, 5) [subsumption resolution 128,49] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
130. 1 | ~2 | ~3 | ~4 | ~5 [avatar contradiction clause 129]
131. $false [avatar sat refutation 50,55,60,65,70,130] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:85,thDist:-85}
% SZS output end Proof for wellord1__t54_wellord1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 6
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 20
% Generated clauses: 55
% Active clauses: 24
% Passive clauses: 24
% Final active clauses: 24
% Main loop iterations started: 24
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 7
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

