i: [1,139,0,2,1,1,0,0] t28_afvect0
a: [3,224,0,2,1,139]
s: 224
i: [1,138,0,2,1,1,0,0] t28_afvect0
a: [3,229,0,2,1,138]
s: 229
i: [1,137,0,2,1,1,0,0] t28_afvect0
a: [3,234,0,2,1,137]
s: 234
i: [1,144,0,10,2,1,0,0] t28_afvect0
a: [3,202,0,4,1,144]
s: 202
i: [1,141,0,4,1,1,0,0] t28_afvect0
a: [3,214,0,4,1,141]
s: 214
i: [1,140,0,4,1,1,0,0] t28_afvect0
a: [3,219,0,4,1,140]
s: 219
a: [3,199,0,6,1,144]
s: 199
i: [1,145,0,2,1,0,0,255] spc1_boole
a: [3,239,0,2,1,145]
s: 239
i: [1,146,0,2,1,0,0,255] fc1_xboole_0
a: [3,244,0,2,1,146]
s: 244
i: [1,186,0,2,1,0,0,2] rc2_xboole_0
a: [3,249,0,2,1,186]
s: 249
i: [1,187,0,2,1,0,0,1] existence_l1_analoaf
a: [3,254,0,2,1,187]
s: 254
i: [1,188,0,2,1,0,0,2] existence_l1_struct_0
a: [3,259,0,2,1,188]
s: 259
i: [1,189,0,2,1,0,0,2] rc1_xboole_0
a: [3,264,0,2,1,189]
s: 264
i: [1,191,0,2,1,0,0,255] rc1_relat_1
a: [3,269,0,2,1,191]
s: 269
i: [1,190,0,2,1,0,0,255] rc1_relat_1
a: [3,274,0,2,1,190]
s: 274
i: [1,192,0,2,1,0,0,255] rc2_relat_1
a: [3,279,0,2,1,192]
s: 279
i: [1,147,0,4,2,0,0,1] dt_l1_analoaf
s: 147
i: [1,162,0,4,2,0,0,255] cc1_relat_1
s: 162
d: [2,281,1,2,1,1,40,147,224]
a: [3,286,1,2,1,281]
s: 286
i: [1,181,0,4,1,0,0,1] existence_m1_subset_1
s: 181
i: [1,163,0,5,2,0,0,255] t6_boole
s: 163
d: [2,282,1,2,1,1,40,147,254]
a: [3,291,1,2,1,282]
s: 291
i: [1,180,0,5,2,0,0,255] t7_boole
s: 180
d: [2,293,1,2,1,1,40,162,244]
a: [3,298,1,2,1,293]
s: 298
i: [1,148,0,6,3,0,0,255] cc7_struct_0
s: 148
i: [1,149,0,6,3,0,0,255] cc5_struct_0
s: 149
d: [2,294,1,2,1,1,40,162,264]
a: [3,303,1,2,1,294]
s: 303
i: [1,151,0,6,3,0,0,255] cc17_struct_0
s: 151
d: [2,308,1,4,2,1,40,148,259]
a: [3,314,1,2,1,308]
s: 314
i: [1,152,0,6,3,0,0,255] cc2_struct_0
s: 152
i: [1,153,0,6,3,0,0,255] cc1_struct_0
s: 153
d: [2,310,2,4,2,1,40,148,291]
a: [3,323,2,2,1,310]
s: 323
i: [1,156,0,6,3,0,0,255] cc4_struct_0
s: 156
i: [1,157,0,6,3,0,0,255] cc6_struct_0
s: 157
i: [1,177,0,6,2,0,0,255] t1_subset
s: 177
i: [1,182,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 182
i: [1,154,0,7,3,0,0,255] cc8_struct_0
s: 154
i: [1,158,0,7,3,0,0,255] cc11_struct_0
s: 158
i: [1,159,0,7,3,0,0,255] cc11_struct_0
s: 159
i: [1,160,0,7,3,0,0,255] cc9_struct_0
s: 160
i: [1,164,0,7,3,0,0,255] fc2_struct_0
s: 164
i: [1,165,0,7,3,0,0,255] fc9_struct_0
s: 165
i: [1,171,0,7,3,0,0,255] fc6_struct_0
s: 171
i: [1,172,0,7,3,0,0,255] fc1_struct_0
s: 172
i: [1,175,0,7,3,0,0,255] fc8_struct_0
s: 175
i: [1,176,0,7,3,0,0,255] fc7_struct_0
s: 176
i: [1,179,0,7,3,0,0,1] t8_boole
s: 179
i: [1,178,0,8,3,0,0,255] t2_subset
s: 178
d: [2,340,1,7,2,1,40,178,214]
a: [3,350,1,3,1,340]
s: 350
r: 350
i: [1,161,0,9,4,0,0,255] cc10_struct_0
s: 161
i: [1,193,0,19,6,0,0,255] t18_afvect0
d: [2,195,0,15,5,0,18,193]
s: 195
a: [3,349,1,3,1,340]
s: 349
i: [1,183,0,18,6,0,0,1] irreflexivity_r1_afvect0
s: 183
d: [2,357,2,4,2,1,40,350,164]
d: [2,359,2,2,1,2,21,357,286]
a: [3,362,2,2,1,359]
s: 362
i: [1,184,0,21,6,0,0,1] dt_k1_afvect0
s: 184
i: [1,185,0,22,7,0,0,1] symmetry_r1_afvect0
s: 185
r: 199
r: 202
i: [1,143,0,9,2,1,0,0] t28_afvect0
a: [3,208,0,3,1,143]
s: 208
a: [3,203,0,4,1,144]
s: 203
a: [3,346,1,4,1,340]
s: 346
a: [3,198,0,6,1,144]
s: 198
i: [1,168,0,23,7,0,0,255] t18_afvect0
s: 168
d: [2,341,1,7,2,1,40,178,219]
a: [3,355,1,4,1,341]
s: 355
i: [1,166,0,26,8,0,0,255] t18_afvect0
s: 166
d: [2,389,2,4,1,1,40,346,182]
a: [3,395,2,4,1,389]
s: 395
i: [1,170,0,29,8,0,0,255] d4_afvect0
s: 170
i: [1,194,0,29,7,0,0,255] d4_afvect0
s: 194
d: [2,398,2,4,1,1,40,355,182]
a: [3,404,2,4,1,398]
s: 404
d: [2,426,0,22,6,0,21,194,184]
s: 426
r: 198
r: 208
a: [3,207,0,3,1,143]
s: 207
r: 395
r: 346
r: 203
r: 214
d: [2,442,0,4,1,2,23,203,207]
a: [3,447,0,4,1,442]
s: 447
s: 199
d: [2,449,0,6,1,2,22,199,207]
a: [3,453,0,6,1,449]
s: 453
d: [2,455,1,26,8,1,44,453,170]
d: [2,456,1,22,7,1,18,455]
d: [2,457,1,18,6,2,21,456,219]
d: [2,458,1,16,5,3,21,457,224]
d: [2,459,1,14,4,4,21,458,229]
d: [2,460,1,12,3,5,21,459,234]
s: 460
d: [2,461,2,9,2,5,47,460]
d: [2,462,2,5,1,5,21,461,219]
a: [3,466,2,5,1,462]
s: 466
d: [2,468,3,0,0,5,52,234,229,224,219,466,195]
e: 468
d: [2,469,3,10,4,1,40,466,195]
d: [2,473,3,6,3,2,21,469,219]
d: [2,474,3,4,2,3,21,473,224]
d: [2,475,3,2,1,4,21,474,229]
d: [2,476,3,0,0,5,21,475,234]
e: 476
r: 207
r: 460
r: 453
r: 466
s: 208
s: 202
s: 214
s: 199
s: 346
a: [3,452,0,6,1,449]
s: 452
d: [2,378,1,18,6,1,40,185,202]
d: [2,379,1,14,5,2,21,378,219]
d: [2,380,1,10,4,3,21,379,214]
d: [2,381,1,8,3,4,21,380,224]
d: [2,382,1,6,2,5,21,381,229]
d: [2,383,1,4,1,6,21,382,234]
a: [3,387,1,4,1,383]
s: 387
s: 395
d: [2,487,1,26,8,1,44,199,170]
d: [2,488,1,22,7,2,21,487,214]
d: [2,489,1,18,6,3,21,488,219]
d: [2,490,1,16,5,4,21,489,224]
d: [2,491,1,14,4,5,21,490,229]
d: [2,492,1,12,3,6,21,491,234]
s: 492
d: [2,427,1,19,6,1,44,426,198]
d: [2,430,1,15,5,2,21,427,214]
d: [2,431,1,11,4,3,21,430,219]
d: [2,432,1,9,3,4,21,431,224]
d: [2,433,1,7,2,5,21,432,229]
d: [2,434,1,5,1,6,21,433,234]
a: [3,437,1,5,1,434]
s: 437
d: [2,517,2,0,0,7,52,234,229,224,214,219,202,437,168]
e: 517
d: [2,518,2,18,6,1,40,437,168]
d: [2,520,2,14,5,2,21,518,202]
d: [2,521,2,10,4,3,21,520,219]
d: [2,522,2,6,3,4,21,521,214]
d: [2,523,2,4,2,5,21,522,224]
d: [2,524,2,2,1,6,21,523,229]
d: [2,525,2,0,0,7,21,524,234]
e: 525
r: 492
r: 199
r: 202
r: 437
r: 387
s: 203
s: 198
a: [3,386,1,4,1,383]
s: 386
a: [3,438,1,5,1,434]
s: 438
d: [2,553,2,0,0,8,52,234,229,224,208,214,219,203,438,166]
e: 553
d: [2,554,2,21,7,1,40,438,166]
d: [2,556,2,18,6,2,21,554,208]
d: [2,557,2,14,5,3,21,556,203]
d: [2,558,2,10,4,4,21,557,219]
d: [2,559,2,6,3,5,21,558,214]
d: [2,560,2,4,2,6,21,559,224]
d: [2,561,2,2,1,7,21,560,229]
d: [2,562,2,0,0,8,21,561,234]
e: 562
i: [1,142,0,13,3,1,0,0] t28_afvect0
d: [2,514,2,9,2,6,47,492]
d: [2,515,2,5,1,6,21,514,214]
d: [2,534,1,19,6,1,44,426,198]
d: [2,542,1,15,5,2,21,534,214]
d: [2,543,1,11,4,3,21,542,219]
d: [2,544,1,9,3,4,21,543,224]
d: [2,545,1,7,2,5,21,544,229]
d: [2,546,1,5,1,6,21,545,234]
f: 562,546,525,515,476,462,449,137,138,139,140,141,142,143,144
% Refutation found. Thanks to Tanya!
% SZS status Theorem for afvect0__t28_afvect0
% SZS output start Proof for afvect0__t28_afvect0
1. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_afvect0(X0,X1,X2) = X2 <=> (r1_afvect0(X0,X2,X1) | X1 = X2))))) [input t28_afvect0]
2. ~! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_afvect0(X0,X1,X2) = X2 <=> (r1_afvect0(X0,X2,X1) | X1 = X2))))) [negated conjecture 1]
18. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (k1_afvect0(X0,X1,X2) = X3 <=> r2_afvect0(X0,X2,X1,X3)))))) [input d4_afvect0]
19. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => m1_subset_1(k1_afvect0(X0,X1,X2),u1_struct_0(X0))) [input dt_k1_afvect0]
42. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_afvect0(X0,X1,X2,X1) <=> (r1_afvect0(X0,X1,X2) | X1 = X2))))) [input t18_afvect0]
55. ? [X0] : (? [X1] : (? [X2] : ((k1_afvect0(X0,X1,X2) = X2 <~> (r1_afvect0(X0,X2,X1) | X1 = X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
56. ? [X0] : (? [X1] : (? [X2] : ((k1_afvect0(X0,X1,X2) = X2 <~> (r1_afvect0(X0,X2,X1) | X1 = X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) [flattening 55]
86. ! [X0] : (! [X1] : (! [X2] : ((r2_afvect0(X0,X1,X2,X1) <=> (r1_afvect0(X0,X1,X2) | X1 = X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 42]
87. ! [X0] : (! [X1] : (! [X2] : ((r2_afvect0(X0,X1,X2,X1) <=> (r1_afvect0(X0,X1,X2) | X1 = X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 86]
88. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((k1_afvect0(X0,X1,X2) = X3 <=> r2_afvect0(X0,X2,X1,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 18]
89. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((k1_afvect0(X0,X1,X2) = X3 <=> r2_afvect0(X0,X2,X1,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 88]
110. ! [X0,X1,X2] : (m1_subset_1(k1_afvect0(X0,X1,X2),u1_struct_0(X0)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 19]
111. ! [X0,X1,X2] : (m1_subset_1(k1_afvect0(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 110]
114. ? [X0] : (? [X1] : (? [X2] : ((((~r1_afvect0(X0,X2,X1) & X1 != X2) | k1_afvect0(X0,X1,X2) != X2) & ((r1_afvect0(X0,X2,X1) | X1 = X2) | k1_afvect0(X0,X1,X2) = X2)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) [nnf transformation 56]
115. ? [X0] : (? [X1] : (? [X2] : (((~r1_afvect0(X0,X2,X1) & X1 != X2) | k1_afvect0(X0,X1,X2) != X2) & (r1_afvect0(X0,X2,X1) | X1 = X2 | k1_afvect0(X0,X1,X2) = X2) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) [flattening 114]
116. ? [X0] : (? [X1] : (? [X2] : (((~r1_afvect0(X0,X2,X1) & X1 != X2) | k1_afvect0(X0,X1,X2) != X2) & (r1_afvect0(X0,X2,X1) | X1 = X2 | k1_afvect0(X0,X1,X2) = X2) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : (((~r1_afvect0(sK0,X2,X1) & X1 != X2) | k1_afvect0(sK0,X1,X2) != X2) & (r1_afvect0(sK0,X2,X1) | X1 = X2 | k1_afvect0(sK0,X1,X2) = X2) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0)) [choice axiom]
117. ? [X1] : (? [X2] : (((~r1_afvect0(sK0,X2,X1) & X1 != X2) | k1_afvect0(sK0,X1,X2) != X2) & (r1_afvect0(sK0,X2,X1) | X1 = X2 | k1_afvect0(sK0,X1,X2) = X2) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (((~r1_afvect0(sK0,X2,sK1) & sK1 != X2) | k1_afvect0(sK0,sK1,X2) != X2) & (r1_afvect0(sK0,X2,sK1) | sK1 = X2 | k1_afvect0(sK0,sK1,X2) = X2) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
118. ? [X2] : (((~r1_afvect0(sK0,X2,sK1) & sK1 != X2) | k1_afvect0(sK0,sK1,X2) != X2) & (r1_afvect0(sK0,X2,sK1) | sK1 = X2 | k1_afvect0(sK0,sK1,X2) = X2) & m1_subset_1(X2,u1_struct_0(sK0))) => (((~r1_afvect0(sK0,sK2,sK1) & sK1 != sK2) | sK2 != k1_afvect0(sK0,sK1,sK2)) & (r1_afvect0(sK0,sK2,sK1) | sK1 = sK2 | sK2 = k1_afvect0(sK0,sK1,sK2)) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
119. ((((~r1_afvect0(sK0,sK2,sK1) & sK1 != sK2) | sK2 != k1_afvect0(sK0,sK1,sK2)) & (r1_afvect0(sK0,sK2,sK1) | sK1 = sK2 | sK2 = k1_afvect0(sK0,sK1,sK2)) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0) [skolemisation 115,118,117,116]
120. ! [X0] : (! [X1] : (! [X2] : (((r2_afvect0(X0,X1,X2,X1) | (~r1_afvect0(X0,X1,X2) & X1 != X2)) & ((r1_afvect0(X0,X1,X2) | X1 = X2) | ~r2_afvect0(X0,X1,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [nnf transformation 87]
121. ! [X0] : (! [X1] : (! [X2] : (((r2_afvect0(X0,X1,X2,X1) | (~r1_afvect0(X0,X1,X2) & X1 != X2)) & (r1_afvect0(X0,X1,X2) | X1 = X2 | ~r2_afvect0(X0,X1,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 120]
122. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((k1_afvect0(X0,X1,X2) = X3 | ~r2_afvect0(X0,X2,X1,X3)) & (r2_afvect0(X0,X2,X1,X3) | k1_afvect0(X0,X1,X2) != X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [nnf transformation 89]
137. ~v7_struct_0(sK0) [cnf transformation 119] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
138. v1_afvect0(sK0) [cnf transformation 119] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
139. l1_analoaf(sK0) [cnf transformation 119] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
140. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 119] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
141. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 119] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
142. r1_afvect0(sK0,sK2,sK1) | sK1 = sK2 | sK2 = k1_afvect0(sK0,sK1,sK2) [cnf transformation 119] {a:0,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
143. sK1 != sK2 | sK2 != k1_afvect0(sK0,sK1,sK2) [cnf transformation 119] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
144. ~r1_afvect0(sK0,sK2,sK1) | sK2 != k1_afvect0(sK0,sK1,sK2) [cnf transformation 119] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
166. ~r2_afvect0(X0,X1,X2,X1) | X1 = X2 | r1_afvect0(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 121] {a:0,w:26,wCS:10400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
167. r2_afvect0(X0,X1,X2,X1) | X1 != X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 121] {a:0,w:22,wCS:8800,sine:255,thAx:0,allAx:1,thDist:-1}
168. r2_afvect0(X0,X1,X2,X1) | ~r1_afvect0(X0,X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 121] {a:0,w:23,wCS:9200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
169. r2_afvect0(X0,X2,X1,X3) | k1_afvect0(X0,X1,X2) != X3 | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 122] {a:0,w:29,wCS:11600,sine:255,thAx:0,allAx:1,thDist:-1}
170. k1_afvect0(X0,X1,X2) = X3 | ~r2_afvect0(X0,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 122] {a:0,w:29,wCS:11600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
184. m1_subset_1(k1_afvect0(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 111] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
193. r2_afvect0(X0,X2,X2,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [equality resolution 167] {a:0,w:19,wCS:7600,sine:255,thAx:0,allAx:1,thDist:-1}
194. r2_afvect0(X0,X2,X1,k1_afvect0(X0,X1,X2)) | ~m1_subset_1(k1_afvect0(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [equality resolution 169] {a:0,w:29,wCS:11600,sine:255,thAx:0,allAx:1,thDist:-1}
195. r2_afvect0(X0,X2,X2,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [duplicate literal removal 193] {a:0,w:15,wCS:6000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
197. 1 <=> sK2 = k1_afvect0(sK0,sK1,sK2) [avatar definition]
198. sK2 = k1_afvect0(sK0,sK1,sK2) <- (1) [avatar component clause 197] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
199. sK2 != k1_afvect0(sK0,sK1,sK2) <- (~1) [avatar component clause 197] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
201. 2 <=> r1_afvect0(sK0,sK2,sK1) [avatar definition]
202. r1_afvect0(sK0,sK2,sK1) <- (2) [avatar component clause 201] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
203. ~r1_afvect0(sK0,sK2,sK1) <- (~2) [avatar component clause 201] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
204. ~1 | ~2 [avatar split clause 144,201,197]
206. 3 <=> sK1 = sK2 [avatar definition]
207. sK1 = sK2 <- (3) [avatar component clause 206] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
208. sK1 != sK2 <- (~3) [avatar component clause 206] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
209. ~1 | ~3 [avatar split clause 143,206,197]
210. 1 | 3 | 2 [avatar split clause 142,201,206,197]
212. 4 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
214. m1_subset_1(sK2,u1_struct_0(sK0)) <- (4) [avatar component clause 212] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
215. 4 [avatar split clause 141,212]
217. 5 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
219. m1_subset_1(sK1,u1_struct_0(sK0)) <- (5) [avatar component clause 217] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
220. 5 [avatar split clause 140,217]
222. 6 <=> l1_analoaf(sK0) [avatar definition]
224. l1_analoaf(sK0) <- (6) [avatar component clause 222] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
225. 6 [avatar split clause 139,222]
227. 7 <=> v1_afvect0(sK0) [avatar definition]
229. v1_afvect0(sK0) <- (7) [avatar component clause 227] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
230. 7 [avatar split clause 138,227]
232. 8 <=> v7_struct_0(sK0) [avatar definition]
234. ~v7_struct_0(sK0) <- (~8) [avatar component clause 232] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
235. ~8 [avatar split clause 137,232]
426. r2_afvect0(X0,X2,X1,k1_afvect0(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [subsumption resolution 194,184] {a:0,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
436. 34 <=> r2_afvect0(sK0,sK2,sK1,sK2) [avatar definition]
437. ~r2_afvect0(sK0,sK2,sK1,sK2) <- (~34) [avatar component clause 436] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
438. r2_afvect0(sK0,sK2,sK1,sK2) <- (34) [avatar component clause 436] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
449. sK1 != k1_afvect0(sK0,sK1,sK1) <- (~1, 3) [forward demodulation 199,207] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
451. 36 <=> sK1 = k1_afvect0(sK0,sK1,sK1) [avatar definition]
453. sK1 != k1_afvect0(sK0,sK1,sK1) <- (~36) [avatar component clause 451] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
454. ~36 | 1 | ~3 [avatar split clause 449,206,197,451]
455. sK1 != X0 | ~r2_afvect0(sK0,sK1,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~36) [superposition 453,170] {a:1,w:26,wCS:2600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
456. sK1 != X0 | ~r2_afvect0(sK0,sK1,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~36) [duplicate literal removal 455] {a:1,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
457. sK1 != X0 | ~r2_afvect0(sK0,sK1,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (5, ~36) [subsumption resolution 456,219] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
458. sK1 != X0 | ~r2_afvect0(sK0,sK1,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (5, 6, ~36) [subsumption resolution 457,224] {a:1,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
459. sK1 != X0 | ~r2_afvect0(sK0,sK1,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | v7_struct_0(sK0) <- (5, 6, 7, ~36) [subsumption resolution 458,229] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
460. sK1 != X0 | ~r2_afvect0(sK0,sK1,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (5, 6, 7, ~8, ~36) [subsumption resolution 459,234] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
461. ~r2_afvect0(sK0,sK1,sK1,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) <- (5, 6, 7, ~8, ~36) [equality resolution 460] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
462. ~r2_afvect0(sK0,sK1,sK1,sK1) <- (5, 6, 7, ~8, ~36) [subsumption resolution 461,219] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
464. 37 <=> r2_afvect0(sK0,sK1,sK1,sK1) [avatar definition]
466. ~r2_afvect0(sK0,sK1,sK1,sK1) <- (~37) [avatar component clause 464] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
467. ~37 | ~5 | ~6 | ~7 | 8 | 36 [avatar split clause 462,451,232,227,222,217,464]
469. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~37) [resolution 466,195] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
473. ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (5, ~37) [subsumption resolution 469,219] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
474. ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (5, 6, ~37) [subsumption resolution 473,224] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
475. v7_struct_0(sK0) <- (5, 6, 7, ~37) [subsumption resolution 474,229] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
476. $false <- (5, 6, 7, ~8, ~37) [subsumption resolution 475,234] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
477. ~5 | ~6 | ~7 | 8 | 37 [avatar contradiction clause 476]
487. sK2 != X0 | ~r2_afvect0(sK0,sK2,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~1) [superposition 199,170] {a:1,w:26,wCS:2600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
488. sK2 != X0 | ~r2_afvect0(sK0,sK2,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~1, 4) [subsumption resolution 487,214] {a:1,w:22,wCS:2200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
489. sK2 != X0 | ~r2_afvect0(sK0,sK2,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~1, 4, 5) [subsumption resolution 488,219] {a:1,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
490. sK2 != X0 | ~r2_afvect0(sK0,sK2,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~1, 4, 5, 6) [subsumption resolution 489,224] {a:1,w:16,wCS:1600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
491. sK2 != X0 | ~r2_afvect0(sK0,sK2,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) | v7_struct_0(sK0) <- (~1, 4, 5, 6, 7) [subsumption resolution 490,229] {a:1,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
492. sK2 != X0 | ~r2_afvect0(sK0,sK2,sK1,X0) | ~m1_subset_1(X0,u1_struct_0(sK0)) <- (~1, 4, 5, 6, 7, ~8) [subsumption resolution 491,234] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
514. ~r2_afvect0(sK0,sK2,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) <- (~1, 4, 5, 6, 7, ~8) [equality resolution 492] {a:2,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
515. ~r2_afvect0(sK0,sK2,sK1,sK2) <- (~1, 4, 5, 6, 7, ~8) [subsumption resolution 514,214] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
516. ~34 | 1 | ~4 | ~5 | ~6 | ~7 | 8 [avatar split clause 515,232,227,222,217,212,197,436]
518. ~r1_afvect0(sK0,sK2,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~34) [resolution 437,168] {a:2,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
520. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (2, ~34) [subsumption resolution 518,202] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
521. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (2, 5, ~34) [subsumption resolution 520,219] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
522. ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (2, 4, 5, ~34) [subsumption resolution 521,214] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
523. ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (2, 4, 5, 6, ~34) [subsumption resolution 522,224] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
524. v7_struct_0(sK0) <- (2, 4, 5, 6, 7, ~34) [subsumption resolution 523,229] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
525. $false <- (2, 4, 5, 6, 7, ~8, ~34) [subsumption resolution 524,234] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
526. ~2 | ~4 | ~5 | ~6 | ~7 | 8 | 34 [avatar contradiction clause 525]
534. r2_afvect0(sK0,sK2,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (1) [superposition 426,198] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
542. r2_afvect0(sK0,sK2,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (1, 4) [subsumption resolution 534,214] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
543. r2_afvect0(sK0,sK2,sK1,sK2) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (1, 4, 5) [subsumption resolution 542,219] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
544. r2_afvect0(sK0,sK2,sK1,sK2) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (1, 4, 5, 6) [subsumption resolution 543,224] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
545. r2_afvect0(sK0,sK2,sK1,sK2) | v7_struct_0(sK0) <- (1, 4, 5, 6, 7) [subsumption resolution 544,229] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
546. r2_afvect0(sK0,sK2,sK1,sK2) <- (1, 4, 5, 6, 7, ~8) [subsumption resolution 545,234] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
547. 34 | ~1 | ~4 | ~5 | ~6 | ~7 | 8 [avatar split clause 546,232,227,222,217,212,197,436]
554. sK1 = sK2 | r1_afvect0(sK0,sK2,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (34) [resolution 438,166] {a:2,w:21,wCS:2100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
556. r1_afvect0(sK0,sK2,sK1) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~3, 34) [subsumption resolution 554,208] {a:2,w:18,wCS:1800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
557. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~2, ~3, 34) [subsumption resolution 556,203] {a:2,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
558. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~2, ~3, 5, 34) [subsumption resolution 557,219] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
559. ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~2, ~3, 4, 5, 34) [subsumption resolution 558,214] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
560. ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~2, ~3, 4, 5, 6, 34) [subsumption resolution 559,224] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
561. v7_struct_0(sK0) <- (~2, ~3, 4, 5, 6, 7, 34) [subsumption resolution 560,229] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
562. $false <- (~2, ~3, 4, 5, 6, 7, ~8, 34) [subsumption resolution 561,234] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
563. 2 | 3 | ~4 | ~5 | ~6 | ~7 | 8 | ~34 [avatar contradiction clause 562]
564. $false [avatar sat refutation 204,209,210,215,220,225,230,235,454,467,477,516,526,547,563] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:379,thDist:-379}
% SZS output end Proof for afvect0__t28_afvect0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 46
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 56
% Generated clauses: 282
% Active clauses: 82
% Passive clauses: 102
% Final active clauses: 64
% Final passive clauses: 12
% Main loop iterations started: 88
% 
% >>> Simplifying Inferences
% Duplicate literals: 38
% Fw subsumption resolutions: 84
% Fw demodulations: 1
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 17
% Equational tautologies: 1
% Forward subsumptions: 18
% 
% >>> Generating Inferences
% Binary resolution: 40
% Unit resulting resolution: 3
% Forward superposition: 10
% Backward superposition: 10
% Self superposition: 2
% Equality resolution: 2
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 17
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 49
% SAT solver unit clauses: 14
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 5245
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

