i: [1,77,0,8,2,1,0,0] t15_mcart_1
a: [3,120,0,4,1,77]
s: 120
i: [1,75,0,7,1,1,0,0] t15_mcart_1
a: [3,130,0,7,1,75]
s: 130
i: [1,78,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,135,0,2,1,78]
s: 135
i: [1,79,0,2,1,0,0,255] fc1_xboole_0
a: [3,140,0,2,1,79]
s: 140
i: [1,103,0,2,1,0,0,2] rc2_xboole_0
a: [3,150,0,2,1,103]
s: 150
i: [1,104,0,2,1,0,0,2] rc1_xboole_0
a: [3,155,0,2,1,104]
s: 155
i: [1,106,0,2,1,0,0,2] rc1_relat_1
a: [3,160,0,2,1,106]
s: 160
i: [1,105,0,2,1,0,0,2] rc1_relat_1
a: [3,165,0,2,1,105]
s: 165
i: [1,107,0,2,1,0,0,255] rc2_relat_1
a: [3,170,0,2,1,107]
s: 170
i: [1,80,0,3,1,0,0,255] d2_xboole_0
a: [3,145,0,3,1,80]
s: 145
r: 135
i: [1,81,0,4,2,0,0,2] cc1_relat_1
s: 81
d: [2,173,1,2,1,1,40,81,155]
a: [3,178,1,2,1,173]
s: 178
i: [1,97,0,4,1,0,0,255] existence_m1_subset_1
s: 97
d: [2,174,1,2,1,1,40,81,140]
a: [3,183,1,2,1,174]
s: 183
i: [1,98,0,4,1,0,0,1] fc3_xboole_0
s: 98
i: [1,99,0,4,1,0,0,1] fc6_relat_1
s: 99
i: [1,82,0,5,2,0,0,255] t6_boole
s: 82
i: [1,88,0,5,2,0,0,1] t7_boole
s: 88
d: [2,188,1,6,1,1,40,88,130]
a: [3,192,1,6,1,188]
s: 192
i: [1,109,0,5,1,0,0,1] d2_tarski
s: 109
i: [1,111,0,5,1,0,0,1] d2_tarski
s: 111
i: [1,86,0,6,2,0,0,255] t1_subset
s: 86
i: [1,101,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 101
d: [2,196,1,7,1,1,40,86,130]
a: [3,202,1,7,1,196]
s: 202
i: [1,87,0,7,3,0,0,1] t8_boole
s: 87
i: [1,100,0,7,1,0,0,1] commutativity_k2_tarski
s: 100
d: [2,204,1,7,1,1,40,101,130]
a: [3,210,1,7,1,204]
s: 210
i: [1,85,0,8,3,0,0,255] t2_subset
s: 85
d: [2,197,1,5,1,0,40,86,109]
s: 197
i: [1,102,0,8,3,0,0,1] fc10_subset_1
s: 102
i: [1,89,0,9,2,0,0,1] t10_mcart_1
s: 89
d: [2,223,1,6,1,1,40,89,130]
a: [3,227,1,6,1,223]
s: 227
i: [1,90,0,9,2,0,0,1] t10_mcart_1
s: 90
d: [2,242,1,0,0,2,52,120,130,90]
e: 242
d: [2,243,1,4,1,1,40,90,130]
d: [2,245,1,0,0,2,21,243,120]
e: 245
r: 120
a: [3,116,0,4,1,77]
s: 116
a: [3,119,0,4,1,77]
s: 119
i: [1,76,0,8,2,1,0,0] t15_mcart_1
a: [3,125,0,4,1,76]
s: 125
d: [2,250,1,2,1,1,40,119,88]
a: [3,264,1,2,1,250]
s: 264
i: [1,112,0,11,3,0,0,1] d2_tarski
s: 112
d: [2,266,2,0,0,3,52,116,125,227,112]
e: 266
d: [2,269,2,8,2,1,40,112,227]
d: [2,273,2,4,1,2,21,269,125]
d: [2,274,2,0,0,3,21,273,116]
e: 274
f: 274,243,223,75,76,77
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t15_mcart_1
% SZS output start Proof for mcart_1__t15_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),X3)) => (r2_hidden(k2_xtuple_0(X0),X3) & (k1_xtuple_0(X0) = X2 | k1_xtuple_0(X0) = X1))) [input t15_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),X3)) => (r2_hidden(k2_xtuple_0(X0),X3) & (k1_xtuple_0(X0) = X2 | k1_xtuple_0(X0) = X1))) [negated conjecture 1]
8. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
26. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [input t10_mcart_1]
32. ~! [X0,X1,X2,X3] : (r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),X3)) => (r2_hidden(k2_xtuple_0(X0),X3) & (k1_xtuple_0(X0) = X2 | k1_xtuple_0(X0) = X1))) [flattening 2]
37. ! [X0,X1,X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [flattening 26]
38. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 8]
42. ? [X0,X1,X2,X3] : ((~r2_hidden(k2_xtuple_0(X0),X3) | (k1_xtuple_0(X0) != X2 & k1_xtuple_0(X0) != X1)) & r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),X3))) [ennf transformation 32]
54. ! [X0,X1,X2] : ((r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1)) | ~r2_hidden(X0,k2_zfmisc_1(X1,X2))) [ennf transformation 37]
58. ? [X0,X1,X2,X3] : ((~r2_hidden(k2_xtuple_0(X0),X3) | (k1_xtuple_0(X0) != X2 & k1_xtuple_0(X0) != X1)) & r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),X3))) => ((~r2_hidden(k2_xtuple_0(sK0),sK3) | (sK2 != k1_xtuple_0(sK0) & sK1 != k1_xtuple_0(sK0))) & r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),sK3))) [choice axiom]
59. (~r2_hidden(k2_xtuple_0(sK0),sK3) | (sK2 != k1_xtuple_0(sK0) & sK1 != k1_xtuple_0(sK0))) & r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),sK3)) [skolemisation 42,58]
60. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 38]
61. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 60]
62. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 61]
63. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
64. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 62,63]
75. r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),sK3)) [cnf transformation 59] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
76. ~r2_hidden(k2_xtuple_0(sK0),sK3) | sK1 != k1_xtuple_0(sK0) [cnf transformation 59] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
77. ~r2_hidden(k2_xtuple_0(sK0),sK3) | sK2 != k1_xtuple_0(sK0) [cnf transformation 59] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
89. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k1_xtuple_0(X0),X1) [cnf transformation 54] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
90. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k2_xtuple_0(X0),X2) [cnf transformation 54] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
91. X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2) | k2_tarski(X0,X1) != X2 [cnf transformation 64] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
112. ~r2_hidden(X4,k2_tarski(X0,X1)) | X0 = X4 | X1 = X4 [equality resolution 91] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
114. 1 <=> sK2 = k1_xtuple_0(sK0) [avatar definition]
116. sK2 != k1_xtuple_0(sK0) <- (~1) [avatar component clause 114] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
118. 2 <=> r2_hidden(k2_xtuple_0(sK0),sK3) [avatar definition]
121. ~1 | ~2 [avatar split clause 77,118,114]
123. 3 <=> sK1 = k1_xtuple_0(sK0) [avatar definition]
125. sK1 != k1_xtuple_0(sK0) <- (~3) [avatar component clause 123] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
126. ~3 | ~2 [avatar split clause 76,118,123]
128. 4 <=> r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),sK3)) [avatar definition]
130. r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),sK3)) <- (4) [avatar component clause 128] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
131. 4 [avatar split clause 75,128]
223. r2_hidden(k1_xtuple_0(sK0),k2_tarski(sK1,sK2)) <- (4) [resolution 89,130] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
225. 18 <=> r2_hidden(k1_xtuple_0(sK0),k2_tarski(sK1,sK2)) [avatar definition]
227. r2_hidden(k1_xtuple_0(sK0),k2_tarski(sK1,sK2)) <- (18) [avatar component clause 225] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
228. 18 | ~4 [avatar split clause 223,128,225]
243. r2_hidden(k2_xtuple_0(sK0),sK3) <- (4) [resolution 90,130] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
247. 2 | ~4 [avatar split clause 243,128,118]
269. sK1 = k1_xtuple_0(sK0) | sK2 = k1_xtuple_0(sK0) <- (18) [resolution 112,227] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
273. sK2 = k1_xtuple_0(sK0) <- (~3, 18) [subsumption resolution 269,125] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
274. $false <- (~1, ~3, 18) [subsumption resolution 273,116] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
275. 1 | 3 | ~18 [avatar contradiction clause 274]
276. $false [avatar sat refutation 121,126,131,228,247,275] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:115,thDist:-115}
% SZS output end Proof for mcart_1__t15_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 30
% >>> Preprocessing
% Introduced skolems: 10
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 33
% Generated clauses: 102
% Active clauses: 38
% Passive clauses: 52
% Final active clauses: 36
% Final passive clauses: 14
% Main loop iterations started: 38
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 2
% Forward subsumptions: 15
% 
% >>> Generating Inferences
% Binary resolution: 25
% Unit resulting resolution: 2
% Forward superposition: 4
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 29
% SAT solver unit clauses: 9
% SAT solver binary clauses: 16
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

