i: [1,29,0,2,1,1,0,0] t37_rewrite1
a: [3,58,0,2,1,29]
s: 58
i: [1,31,0,4,1,1,0,0] t37_rewrite1
a: [3,44,0,4,1,31]
s: 44
i: [1,30,0,8,2,1,0,0] t37_rewrite1
a: [3,53,0,4,1,30]
s: 53
i: [1,32,0,10,3,0,0,1] t25_rewrite1
s: 32
i: [1,33,0,10,3,0,0,1] t25_rewrite1
s: 33
i: [1,34,0,13,3,0,0,1] d7_rewrite1
s: 34
i: [1,35,0,13,3,0,0,1] d7_rewrite1
s: 35
i: [1,37,0,13,3,0,0,1] d8_rewrite1
s: 37
i: [1,38,0,13,3,0,0,1] d8_rewrite1
s: 38
i: [1,36,0,14,4,0,0,1] d7_rewrite1
s: 36
i: [1,39,0,14,4,0,0,1] d8_rewrite1
s: 39
d: [2,61,1,15,4,0,40,34,32]
d: [2,62,1,13,3,0,18,61]
s: 62
i: [1,40,0,14,4,0,0,1] t30_rewrite1
s: 40
d: [2,60,1,15,4,0,40,34,33]
d: [2,63,1,13,3,0,18,60]
s: 63
d: [2,65,1,15,4,0,40,35,32]
d: [2,66,1,13,3,0,18,65]
s: 66
d: [2,64,1,15,4,0,40,35,33]
d: [2,67,1,13,3,0,18,64]
s: 67
d: [2,69,1,15,4,0,40,37,32]
d: [2,70,1,13,3,0,18,69]
s: 70
d: [2,68,1,15,4,0,40,37,33]
d: [2,71,1,13,3,0,18,68]
s: 71
d: [2,73,1,15,4,0,40,38,32]
d: [2,74,1,13,3,0,18,73]
s: 74
d: [2,72,1,15,4,0,40,38,33]
d: [2,75,1,13,3,0,18,72]
s: 75
d: [2,79,1,19,5,0,40,36,38]
d: [2,80,1,17,4,0,18,79]
s: 80
d: [2,78,1,19,5,0,40,36,37]
d: [2,81,1,17,4,0,18,78]
s: 81
d: [2,77,1,19,5,0,40,36,35]
d: [2,82,1,17,4,0,18,77]
s: 82
d: [2,109,2,16,5,0,40,82,35]
d: [2,114,2,10,3,0,18,109]
s: 114
d: [2,76,1,19,5,0,40,36,34]
d: [2,83,1,17,4,0,18,76]
s: 83
d: [2,87,1,19,5,0,40,39,38]
d: [2,88,1,17,4,0,18,87]
s: 88
d: [2,121,2,16,5,0,40,83,35]
d: [2,126,2,10,3,0,18,121]
s: 126
d: [2,86,1,19,5,0,40,39,37]
d: [2,89,1,17,4,0,18,86]
s: 89
d: [2,120,2,16,5,0,40,83,34]
d: [2,127,2,10,3,0,18,120]
s: 127
d: [2,85,1,19,5,0,40,39,35]
d: [2,90,1,17,4,0,18,85]
s: 90
d: [2,84,1,19,5,0,40,39,34]
d: [2,91,1,17,4,0,18,84]
s: 91
d: [2,131,2,16,5,0,40,88,38]
d: [2,132,2,10,3,0,18,131]
s: 132
d: [2,143,2,16,5,0,40,89,38]
d: [2,144,2,10,3,0,18,143]
s: 144
d: [2,152,3,6,2,1,40,132,53]
d: [2,157,3,4,1,2,21,152,58]
a: [3,161,3,4,1,157]
s: 161
d: [2,142,2,16,5,0,40,89,37]
d: [2,145,2,10,3,0,18,142]
s: 145
r: 53
a: [3,49,0,4,1,30]
s: 49
r: 49
a: [3,48,0,4,1,30]
s: 48
s: 53
d: [2,163,3,6,2,1,40,144,53]
d: [2,168,3,4,1,2,21,163,58]
a: [3,172,3,4,1,168]
s: 172
d: [2,92,2,19,5,0,40,40,62]
d: [2,93,2,17,4,0,18,92]
s: 93
d: [2,94,2,19,5,0,40,63,40]
d: [2,95,2,17,4,0,18,94]
s: 95
d: [2,176,3,6,2,1,40,145,53]
d: [2,182,3,4,1,2,21,176,58]
a: [3,186,3,4,1,182]
s: 186
d: [2,96,2,19,5,0,40,66,40]
d: [2,97,2,17,4,0,18,96]
s: 97
d: [2,224,3,16,5,0,40,95,66]
d: [2,235,3,10,3,0,18,224]
s: 235
d: [2,98,2,19,5,0,40,67,40]
d: [2,99,2,17,4,0,18,98]
s: 99
d: [2,100,2,19,5,0,40,70,40]
d: [2,101,2,17,4,0,18,100]
s: 101
d: [2,223,3,16,5,0,40,95,62]
d: [2,236,3,10,3,0,18,223]
s: 236
d: [2,102,2,19,5,0,40,71,40]
d: [2,103,2,17,4,0,18,102]
s: 103
d: [2,249,3,16,5,0,40,99,66]
d: [2,262,3,10,3,0,18,249]
s: 262
d: [2,104,2,19,5,0,40,74,40]
d: [2,105,2,17,4,0,18,104]
s: 105
d: [2,106,2,19,5,0,40,75,40]
d: [2,107,2,17,4,0,18,106]
s: 107
d: [2,248,3,16,5,0,40,99,62]
d: [2,263,3,10,3,0,18,248]
s: 263
d: [2,111,2,25,5,0,40,82,38]
d: [2,112,2,23,4,0,18,111]
s: 112
d: [2,265,3,16,5,0,40,101,75]
d: [2,274,3,10,3,0,18,265]
s: 274
d: [2,326,4,0,0,3,52,58,44,172,274]
e: 326
d: [2,329,4,6,2,1,40,274,172]
d: [2,338,4,2,1,2,21,329,44]
d: [2,339,4,0,0,3,21,338,58]
e: 339
r: 48
r: 53
r: 172
r: 161
s: 49
d: [2,352,4,0,0,3,52,58,44,49,263]
e: 352
d: [2,359,4,6,2,1,40,49,263]
d: [2,367,4,2,1,2,21,359,44]
d: [2,368,4,0,0,3,21,367,58]
e: 368
d: [2,210,3,6,2,1,40,53,144]
d: [2,214,3,4,1,2,21,210,58]
f: 368,339,214,29,30,31
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rewrite1__t37_rewrite1
% SZS output start Proof for rewrite1__t37_rewrite1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ((r6_rewrite1(X0,X1,X2) | r5_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X2))) [input t37_rewrite1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ((r6_rewrite1(X0,X1,X2) | r5_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X2))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r5_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3)))) [input d7_rewrite1]
4. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r6_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1)))) [input d8_rewrite1]
5. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r1_rewrite1(X0,X1,X2) => (r2_rewrite1(X0,X2,X1) & r2_rewrite1(X0,X1,X2)))) [input t25_rewrite1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : ! [X3] : ((r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X3))) [input t30_rewrite1]
7. ~! [X0] : (v1_relat_1(X0) => ! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | r5_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X2))) [flattening 2]
8. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r1_rewrite1(X0,X1,X2) => (r2_rewrite1(X0,X2,X1) & r2_rewrite1(X0,X1,X2)))) [flattening 5]
9. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r5_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3)))) [flattening 3]
10. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r6_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1)))) [flattening 4]
11. ! [X0] : (v1_relat_1(X0) => ! [X1,X2,X3] : ((r2_rewrite1(X0,X2,X3) & r2_rewrite1(X0,X1,X2)) => r2_rewrite1(X0,X1,X3))) [flattening 6]
12. ? [X0] : (? [X1,X2] : (~r2_rewrite1(X0,X1,X2) & (r6_rewrite1(X0,X1,X2) | r5_rewrite1(X0,X1,X2))) & v1_relat_1(X0)) [ennf transformation 7]
13. ! [X0] : (! [X1,X2] : ((r2_rewrite1(X0,X2,X1) & r2_rewrite1(X0,X1,X2)) | ~r1_rewrite1(X0,X1,X2)) | ~v1_relat_1(X0)) [ennf transformation 8]
14. ! [X0] : (! [X1,X2] : (r5_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3))) | ~v1_relat_1(X0)) [ennf transformation 9]
15. ! [X0] : (! [X1,X2] : (r6_rewrite1(X0,X1,X2) <=> ? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1))) | ~v1_relat_1(X0)) [ennf transformation 10]
16. ! [X0] : (! [X1,X2,X3] : (r2_rewrite1(X0,X1,X3) | (~r2_rewrite1(X0,X2,X3) | ~r2_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [ennf transformation 11]
17. ! [X0] : (! [X1,X2,X3] : (r2_rewrite1(X0,X1,X3) | ~r2_rewrite1(X0,X2,X3) | ~r2_rewrite1(X0,X1,X2)) | ~v1_relat_1(X0)) [flattening 16]
18. ? [X0] : (? [X1,X2] : (~r2_rewrite1(X0,X1,X2) & (r6_rewrite1(X0,X1,X2) | r5_rewrite1(X0,X1,X2))) & v1_relat_1(X0)) => (? [X2,X1] : (~r2_rewrite1(sK0,X1,X2) & (r6_rewrite1(sK0,X1,X2) | r5_rewrite1(sK0,X1,X2))) & v1_relat_1(sK0)) [choice axiom]
19. ? [X2,X1] : (~r2_rewrite1(sK0,X1,X2) & (r6_rewrite1(sK0,X1,X2) | r5_rewrite1(sK0,X1,X2))) => (~r2_rewrite1(sK0,sK1,sK2) & (r6_rewrite1(sK0,sK1,sK2) | r5_rewrite1(sK0,sK1,sK2))) [choice axiom]
20. (~r2_rewrite1(sK0,sK1,sK2) & (r6_rewrite1(sK0,sK1,sK2) | r5_rewrite1(sK0,sK1,sK2))) & v1_relat_1(sK0) [skolemisation 12,19,18]
21. ! [X0] : (! [X1,X2] : ((r5_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X3))) & (? [X3] : (r1_rewrite1(X0,X2,X3) & r1_rewrite1(X0,X1,X3)) | ~r5_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [nnf transformation 14]
22. ! [X0] : (! [X1,X2] : ((r5_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X3))) & (? [X4] : (r1_rewrite1(X0,X2,X4) & r1_rewrite1(X0,X1,X4)) | ~r5_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [rectify 21]
23. ! [X2,X1,X0] : (? [X4] : (r1_rewrite1(X0,X2,X4) & r1_rewrite1(X0,X1,X4)) => (r1_rewrite1(X0,X2,sK3(X0,X1,X2)) & r1_rewrite1(X0,X1,sK3(X0,X1,X2)))) [choice axiom]
24. ! [X0] : (! [X1,X2] : ((r5_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X2,X3) | ~r1_rewrite1(X0,X1,X3))) & ((r1_rewrite1(X0,X2,sK3(X0,X1,X2)) & r1_rewrite1(X0,X1,sK3(X0,X1,X2))) | ~r5_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [skolemisation 22,23]
25. ! [X0] : (! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X3,X2) | ~r1_rewrite1(X0,X3,X1))) & (? [X3] : (r1_rewrite1(X0,X3,X2) & r1_rewrite1(X0,X3,X1)) | ~r6_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [nnf transformation 15]
26. ! [X0] : (! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X3,X2) | ~r1_rewrite1(X0,X3,X1))) & (? [X4] : (r1_rewrite1(X0,X4,X2) & r1_rewrite1(X0,X4,X1)) | ~r6_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [rectify 25]
27. ! [X2,X1,X0] : (? [X4] : (r1_rewrite1(X0,X4,X2) & r1_rewrite1(X0,X4,X1)) => (r1_rewrite1(X0,sK4(X0,X1,X2),X2) & r1_rewrite1(X0,sK4(X0,X1,X2),X1))) [choice axiom]
28. ! [X0] : (! [X1,X2] : ((r6_rewrite1(X0,X1,X2) | ! [X3] : (~r1_rewrite1(X0,X3,X2) | ~r1_rewrite1(X0,X3,X1))) & ((r1_rewrite1(X0,sK4(X0,X1,X2),X2) & r1_rewrite1(X0,sK4(X0,X1,X2),X1)) | ~r6_rewrite1(X0,X1,X2))) | ~v1_relat_1(X0)) [skolemisation 26,27]
29. v1_relat_1(sK0) [cnf transformation 20] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
30. r6_rewrite1(sK0,sK1,sK2) | r5_rewrite1(sK0,sK1,sK2) [cnf transformation 20] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
31. ~r2_rewrite1(sK0,sK1,sK2) [cnf transformation 20] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
32. ~r1_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 13] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. ~r1_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X2,X1) | ~v1_relat_1(X0) [cnf transformation 13] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. r1_rewrite1(X0,X1,sK3(X0,X1,X2)) | ~r5_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
35. r1_rewrite1(X0,X2,sK3(X0,X1,X2)) | ~r5_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
37. r1_rewrite1(X0,sK4(X0,X1,X2),X1) | ~r6_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
38. r1_rewrite1(X0,sK4(X0,X1,X2),X2) | ~r6_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
39. ~r1_rewrite1(X0,X3,X2) | r6_rewrite1(X0,X1,X2) | ~r1_rewrite1(X0,X3,X1) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
40. ~r2_rewrite1(X0,X2,X3) | r2_rewrite1(X0,X1,X3) | ~r2_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [cnf transformation 17] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. 1 <=> r2_rewrite1(sK0,sK1,sK2) [avatar definition]
44. ~r2_rewrite1(sK0,sK1,sK2) <- (~1) [avatar component clause 42] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
45. ~1 [avatar split clause 31,42]
47. 2 <=> r5_rewrite1(sK0,sK1,sK2) [avatar definition]
49. r5_rewrite1(sK0,sK1,sK2) <- (2) [avatar component clause 47] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
51. 3 <=> r6_rewrite1(sK0,sK1,sK2) [avatar definition]
53. r6_rewrite1(sK0,sK1,sK2) <- (3) [avatar component clause 51] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
54. 2 | 3 [avatar split clause 30,51,47]
56. 4 <=> v1_relat_1(sK0) [avatar definition]
58. v1_relat_1(sK0) <- (4) [avatar component clause 56] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
59. 4 [avatar split clause 29,56]
61. ~r5_rewrite1(X3,X4,X5) | ~v1_relat_1(X3) | r2_rewrite1(X3,X4,sK3(X3,X4,X5)) | ~v1_relat_1(X3) [resolution 34,32] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:3,thDist:-3}
62. r2_rewrite1(X3,X4,sK3(X3,X4,X5)) | ~v1_relat_1(X3) | ~r5_rewrite1(X3,X4,X5) [duplicate literal removal 61] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
64. ~r5_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) | r2_rewrite1(X0,sK3(X0,X1,X2),X2) | ~v1_relat_1(X0) [resolution 35,33] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:3,thDist:-3}
67. r2_rewrite1(X0,sK3(X0,X1,X2),X2) | ~v1_relat_1(X0) | ~r5_rewrite1(X0,X1,X2) [duplicate literal removal 64] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
69. ~r6_rewrite1(X3,X4,X5) | ~v1_relat_1(X3) | r2_rewrite1(X3,sK4(X3,X4,X5),X4) | ~v1_relat_1(X3) [resolution 37,32] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:3,thDist:-3}
70. r2_rewrite1(X3,sK4(X3,X4,X5),X4) | ~v1_relat_1(X3) | ~r6_rewrite1(X3,X4,X5) [duplicate literal removal 69] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
72. ~r6_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) | r2_rewrite1(X0,X2,sK4(X0,X1,X2)) | ~v1_relat_1(X0) [resolution 38,33] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:3,thDist:-3}
75. r2_rewrite1(X0,X2,sK4(X0,X1,X2)) | ~v1_relat_1(X0) | ~r6_rewrite1(X0,X1,X2) [duplicate literal removal 72] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
86. r6_rewrite1(X8,X9,X10) | ~r1_rewrite1(X8,sK4(X8,X10,X11),X9) | ~v1_relat_1(X8) | ~r6_rewrite1(X8,X10,X11) | ~v1_relat_1(X8) [resolution 39,37] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
89. ~r1_rewrite1(X8,sK4(X8,X10,X11),X9) | r6_rewrite1(X8,X9,X10) | ~v1_relat_1(X8) | ~r6_rewrite1(X8,X10,X11) [duplicate literal removal 86] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
98. ~v1_relat_1(X0) | ~r5_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X3,X2) | ~r2_rewrite1(X0,X3,sK3(X0,X1,X2)) | ~v1_relat_1(X0) [resolution 67,40] {a:2,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
99. ~r2_rewrite1(X0,X3,sK3(X0,X1,X2)) | ~r5_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X3,X2) | ~v1_relat_1(X0) [duplicate literal removal 98] {a:2,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
100. ~v1_relat_1(X0) | ~r6_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X3,X1) | ~r2_rewrite1(X0,X3,sK4(X0,X1,X2)) | ~v1_relat_1(X0) [resolution 70,40] {a:2,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
101. ~r2_rewrite1(X0,X3,sK4(X0,X1,X2)) | ~r6_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X3,X1) | ~v1_relat_1(X0) [duplicate literal removal 100] {a:2,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
143. r6_rewrite1(X11,X12,X13) | ~v1_relat_1(X11) | ~r6_rewrite1(X11,X13,X12) | ~r6_rewrite1(X11,X13,X12) | ~v1_relat_1(X11) [resolution 89,38] {a:2,w:16,wCS:6400,sine:1,thAx:0,allAx:6,thDist:-6}
144. ~r6_rewrite1(X11,X13,X12) | ~v1_relat_1(X11) | r6_rewrite1(X11,X12,X13) [duplicate literal removal 143] {a:2,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
170. 6 <=> r6_rewrite1(sK0,sK2,sK1) [avatar definition]
172. r6_rewrite1(sK0,sK2,sK1) <- (6) [avatar component clause 170] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
210. ~v1_relat_1(sK0) | r6_rewrite1(sK0,sK2,sK1) <- (3) [resolution 53,144] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
214. r6_rewrite1(sK0,sK2,sK1) <- (3, 4) [subsumption resolution 210,58] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
215. 6 | ~3 | ~4 [avatar split clause 214,56,51,170]
248. ~r5_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) | ~v1_relat_1(X0) | ~r5_rewrite1(X0,X1,X2) [resolution 99,62] {a:3,w:16,wCS:6400,sine:1,thAx:0,allAx:7,thDist:-7}
263. ~r5_rewrite1(X0,X1,X2) | r2_rewrite1(X0,X1,X2) | ~v1_relat_1(X0) [duplicate literal removal 248] {a:3,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:7,thDist:-7}
265. ~r6_rewrite1(X3,X4,X5) | r2_rewrite1(X3,X5,X4) | ~v1_relat_1(X3) | ~v1_relat_1(X3) | ~r6_rewrite1(X3,X4,X5) [resolution 101,75] {a:3,w:16,wCS:6400,sine:1,thAx:0,allAx:7,thDist:-7}
274. ~r6_rewrite1(X3,X4,X5) | r2_rewrite1(X3,X5,X4) | ~v1_relat_1(X3) [duplicate literal removal 265] {a:3,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:7,thDist:-7}
329. r2_rewrite1(sK0,sK1,sK2) | ~v1_relat_1(sK0) <- (6) [resolution 274,172] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:16,thDist:-16}
338. ~v1_relat_1(sK0) <- (~1, 6) [subsumption resolution 329,44] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:16,thDist:-16}
339. $false <- (~1, 4, 6) [subsumption resolution 338,58] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:16,thDist:-16}
340. 1 | ~4 | ~6 [avatar contradiction clause 339]
359. r2_rewrite1(sK0,sK1,sK2) | ~v1_relat_1(sK0) <- (2) [resolution 49,263] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
367. ~v1_relat_1(sK0) <- (~1, 2) [subsumption resolution 359,44] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
368. $false <- (~1, 2, 4) [subsumption resolution 367,58] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
369. 1 | ~2 | ~4 [avatar contradiction clause 368]
370. $false [avatar sat refutation 45,54,59,215,340,369] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:229,thDist:-229}
% SZS output end Proof for rewrite1__t37_rewrite1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 5
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 297
% Active clauses: 55
% Passive clauses: 124
% Final active clauses: 49
% Final passive clauses: 68
% Main loop iterations started: 55
% 
% >>> Simplifying Inferences
% Duplicate literals: 129
% Fw subsumption resolutions: 22
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 134
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 2
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% resolution: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

