i: [1,54,0,2,1,1,0,0] t14_e_siec
a: [3,98,0,2,1,54]
s: 98
i: [1,53,0,2,1,1,0,0] t14_e_siec
a: [3,103,0,2,1,53]
s: 103
i: [1,55,0,24,4,1,0,0] t14_e_siec
a: [3,81,0,6,1,55]
s: 81
i: [1,74,0,2,1,0,0,1] existence_l1_e_siec
a: [3,108,0,2,1,74]
s: 108
i: [1,75,0,2,1,0,0,1] rc3_e_siec
a: [3,113,0,2,1,75]
s: 113
i: [1,77,0,2,1,0,0,1] rc2_e_siec
a: [3,118,0,2,1,77]
s: 118
i: [1,76,0,2,1,0,0,1] rc2_e_siec
a: [3,123,0,2,1,76]
s: 123
i: [1,71,0,3,1,0,0,1] reflexivity_r1_tarski
s: 71
i: [1,70,0,4,1,0,0,255] existence_m1_subset_1
s: 70
i: [1,72,0,4,1,0,0,2] fc6_relat_1
s: 72
i: [1,73,0,4,1,0,0,255] dt_k3_relat_1
s: 73
i: [1,56,0,5,2,0,0,1] dt_u2_e_siec
s: 56
i: [1,57,0,5,2,0,0,1] dt_u1_e_siec
s: 57
i: [1,66,0,7,2,0,0,255] t3_subset
s: 66
i: [1,67,0,7,2,0,0,255] t3_subset
s: 67
i: [1,65,0,8,3,0,0,255] cc2_relat_1
s: 65
d: [2,125,1,5,1,0,40,66,70]
s: 125
i: [1,68,0,11,3,0,0,1] t3_sysrel
s: 68
d: [2,127,1,6,2,0,40,65,70]
s: 127
i: [1,69,0,11,3,0,0,1] t3_sysrel
s: 69
i: [1,58,0,12,3,0,0,255] d1_e_siec
s: 58
d: [2,129,1,10,2,0,40,68,71]
d: [2,131,1,6,1,0,21,129,72]
s: 131
i: [1,59,0,12,3,0,0,255] d1_e_siec
s: 59
d: [2,132,1,10,2,0,40,69,71]
d: [2,134,1,6,1,0,21,132,72]
s: 134
i: [1,60,0,12,3,0,0,255] d1_e_siec
s: 60
i: [1,61,0,12,3,0,0,255] d1_e_siec
s: 61
d: [2,128,1,7,3,0,40,65,67]
s: 128
i: [1,62,0,12,3,0,0,255] d1_e_siec
s: 62
d: [2,151,2,7,2,0,40,128,131]
s: 151
i: [1,63,0,12,3,0,0,255] d1_e_siec
s: 63
i: [1,64,0,52,8,0,0,255] d1_e_siec
s: 64
d: [2,154,2,7,2,0,40,128,134]
s: 154
d: [2,135,1,13,4,0,40,58,69]
d: [2,137,1,10,3,0,21,135,57]
s: 137
d: [2,136,1,13,4,0,40,58,68]
d: [2,138,1,10,3,0,21,136,57]
s: 138
d: [2,160,2,0,0,3,52,103,98,81,138]
e: 160
d: [2,161,2,4,2,1,40,138,81]
d: [2,164,2,2,1,2,21,161,98]
d: [2,165,2,0,0,3,21,164,103]
e: 165
r: 81
a: [3,80,0,6,1,55]
s: 80
a: [3,85,0,6,1,55]
s: 85
d: [2,177,2,0,0,3,52,103,98,85,137]
e: 177
d: [2,178,2,4,2,1,40,85,137]
d: [2,180,2,2,1,2,21,178,98]
d: [2,181,2,0,0,3,21,180,103]
e: 181
r: 85
d: [2,167,2,7,2,1,40,80,128]
a: [3,175,2,3,1,167]
s: 175
a: [3,84,0,6,1,55]
s: 84
a: [3,89,0,6,1,55]
s: 89
d: [2,141,1,13,4,0,40,59,69]
d: [2,143,1,10,3,0,21,141,56]
s: 143
d: [2,142,1,13,4,0,40,59,68]
d: [2,144,1,10,3,0,21,142,56]
s: 144
d: [2,185,2,0,0,3,52,103,98,89,144]
e: 185
d: [2,186,2,4,2,1,40,144,89]
d: [2,189,2,2,1,2,21,186,98]
d: [2,190,2,0,0,3,21,189,103]
e: 190
r: 89
a: [3,88,0,6,1,55]
s: 88
a: [3,93,0,6,1,55]
s: 93
d: [2,193,2,0,0,3,52,103,98,93,143]
e: 193
d: [2,194,2,4,2,1,40,93,143]
d: [2,196,2,2,1,2,21,194,98]
d: [2,197,2,0,0,3,21,196,103]
e: 197
f: 197,190,181,165,53,54,55
% Refutation found. Thanks to Tanya!
% SZS status Theorem for e_siec__t14_e_siec
% SZS output start Proof for e_siec__t14_e_siec
1. ! [X0] : ((l1_e_siec(X0) & v3_e_siec(X0) & v2_e_siec(X0)) => (r1_tarski(k10_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) & r1_tarski(k9_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) & r1_tarski(k10_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)) & r1_tarski(k9_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)))) [input t14_e_siec]
2. ~! [X0] : ((l1_e_siec(X0) & v3_e_siec(X0) & v2_e_siec(X0)) => (r1_tarski(k10_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) & r1_tarski(k9_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) & r1_tarski(k10_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)) & r1_tarski(k9_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)))) [negated conjecture 1]
4. ! [X0] : (l1_e_siec(X0) => (v2_e_siec(X0) <=> (u2_e_siec(X0) = k3_relat_1(u2_e_siec(X0),u1_e_siec(X0)) & u2_e_siec(X0) = k3_relat_1(u2_e_siec(X0),u2_e_siec(X0)) & u1_e_siec(X0) = k3_relat_1(u1_e_siec(X0),u2_e_siec(X0)) & u1_e_siec(X0) = k3_relat_1(u1_e_siec(X0),u1_e_siec(X0)) & r1_tarski(u2_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))) & r1_tarski(u1_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0)))))) [input d1_e_siec]
13. ! [X0] : (l1_e_siec(X0) => v1_relat_1(u1_e_siec(X0))) [input dt_u1_e_siec]
15. ! [X0] : (l1_e_siec(X0) => v1_relat_1(u2_e_siec(X0))) [input dt_u2_e_siec]
24. ! [X0] : ! [X1] : ! [X2] : (v1_relat_1(X2) => (r1_tarski(X2,k2_zfmisc_1(X0,X1)) => (r1_tarski(k10_xtuple_0(X2),X1) & r1_tarski(k9_xtuple_0(X2),X0)))) [input t3_sysrel]
26. ! [X0,X1,X2] : (v1_relat_1(X2) => (r1_tarski(X2,k2_zfmisc_1(X0,X1)) => (r1_tarski(k10_xtuple_0(X2),X1) & r1_tarski(k9_xtuple_0(X2),X0)))) [flattening 24]
30. ~! [X0] : ((l1_e_siec(X0) & v2_e_siec(X0)) => (r1_tarski(k10_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) & r1_tarski(k9_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) & r1_tarski(k10_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)) & r1_tarski(k9_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)))) [pure predicate removal 2]
32. ? [X0] : ((~r1_tarski(k10_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) | ~r1_tarski(k9_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) | ~r1_tarski(k10_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)) | ~r1_tarski(k9_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0))) & (l1_e_siec(X0) & v2_e_siec(X0))) [ennf transformation 30]
33. ? [X0] : ((~r1_tarski(k10_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) | ~r1_tarski(k9_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) | ~r1_tarski(k10_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)) | ~r1_tarski(k9_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0))) & l1_e_siec(X0) & v2_e_siec(X0)) [flattening 32]
34. ! [X0] : (v1_relat_1(u2_e_siec(X0)) | ~l1_e_siec(X0)) [ennf transformation 15]
35. ! [X0] : (v1_relat_1(u1_e_siec(X0)) | ~l1_e_siec(X0)) [ennf transformation 13]
36. ! [X0] : ((v2_e_siec(X0) <=> (u2_e_siec(X0) = k3_relat_1(u2_e_siec(X0),u1_e_siec(X0)) & u2_e_siec(X0) = k3_relat_1(u2_e_siec(X0),u2_e_siec(X0)) & u1_e_siec(X0) = k3_relat_1(u1_e_siec(X0),u2_e_siec(X0)) & u1_e_siec(X0) = k3_relat_1(u1_e_siec(X0),u1_e_siec(X0)) & r1_tarski(u2_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))) & r1_tarski(u1_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))))) | ~l1_e_siec(X0)) [ennf transformation 4]
38. ! [X0,X1,X2] : (((r1_tarski(k10_xtuple_0(X2),X1) & r1_tarski(k9_xtuple_0(X2),X0)) | ~r1_tarski(X2,k2_zfmisc_1(X0,X1))) | ~v1_relat_1(X2)) [ennf transformation 26]
39. ! [X0,X1,X2] : ((r1_tarski(k10_xtuple_0(X2),X1) & r1_tarski(k9_xtuple_0(X2),X0)) | ~r1_tarski(X2,k2_zfmisc_1(X0,X1)) | ~v1_relat_1(X2)) [flattening 38]
40. ? [X0] : ((~r1_tarski(k10_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) | ~r1_tarski(k9_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) | ~r1_tarski(k10_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)) | ~r1_tarski(k9_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0))) & l1_e_siec(X0) & v2_e_siec(X0)) => ((~r1_tarski(k10_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) | ~r1_tarski(k9_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) | ~r1_tarski(k10_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0)) | ~r1_tarski(k9_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0))) & l1_e_siec(sK0) & v2_e_siec(sK0)) [choice axiom]
41. (~r1_tarski(k10_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) | ~r1_tarski(k9_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) | ~r1_tarski(k10_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0)) | ~r1_tarski(k9_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0))) & l1_e_siec(sK0) & v2_e_siec(sK0) [skolemisation 33,40]
42. ! [X0] : (((v2_e_siec(X0) | (u2_e_siec(X0) != k3_relat_1(u2_e_siec(X0),u1_e_siec(X0)) | u2_e_siec(X0) != k3_relat_1(u2_e_siec(X0),u2_e_siec(X0)) | u1_e_siec(X0) != k3_relat_1(u1_e_siec(X0),u2_e_siec(X0)) | u1_e_siec(X0) != k3_relat_1(u1_e_siec(X0),u1_e_siec(X0)) | ~r1_tarski(u2_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))) | ~r1_tarski(u1_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))))) & ((u2_e_siec(X0) = k3_relat_1(u2_e_siec(X0),u1_e_siec(X0)) & u2_e_siec(X0) = k3_relat_1(u2_e_siec(X0),u2_e_siec(X0)) & u1_e_siec(X0) = k3_relat_1(u1_e_siec(X0),u2_e_siec(X0)) & u1_e_siec(X0) = k3_relat_1(u1_e_siec(X0),u1_e_siec(X0)) & r1_tarski(u2_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))) & r1_tarski(u1_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0)))) | ~v2_e_siec(X0))) | ~l1_e_siec(X0)) [nnf transformation 36]
43. ! [X0] : (((v2_e_siec(X0) | u2_e_siec(X0) != k3_relat_1(u2_e_siec(X0),u1_e_siec(X0)) | u2_e_siec(X0) != k3_relat_1(u2_e_siec(X0),u2_e_siec(X0)) | u1_e_siec(X0) != k3_relat_1(u1_e_siec(X0),u2_e_siec(X0)) | u1_e_siec(X0) != k3_relat_1(u1_e_siec(X0),u1_e_siec(X0)) | ~r1_tarski(u2_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))) | ~r1_tarski(u1_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0)))) & ((u2_e_siec(X0) = k3_relat_1(u2_e_siec(X0),u1_e_siec(X0)) & u2_e_siec(X0) = k3_relat_1(u2_e_siec(X0),u2_e_siec(X0)) & u1_e_siec(X0) = k3_relat_1(u1_e_siec(X0),u2_e_siec(X0)) & u1_e_siec(X0) = k3_relat_1(u1_e_siec(X0),u1_e_siec(X0)) & r1_tarski(u2_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))) & r1_tarski(u1_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0)))) | ~v2_e_siec(X0))) | ~l1_e_siec(X0)) [flattening 42]
53. v2_e_siec(sK0) [cnf transformation 41] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
54. l1_e_siec(sK0) [cnf transformation 41] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. ~r1_tarski(k10_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) | ~r1_tarski(k9_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) | ~r1_tarski(k10_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0)) | ~r1_tarski(k9_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0)) [cnf transformation 41] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
56. v1_relat_1(u2_e_siec(X0)) | ~l1_e_siec(X0) [cnf transformation 34] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. v1_relat_1(u1_e_siec(X0)) | ~l1_e_siec(X0) [cnf transformation 35] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. r1_tarski(u1_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))) | ~v2_e_siec(X0) | ~l1_e_siec(X0) [cnf transformation 43] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
59. r1_tarski(u2_e_siec(X0),k2_zfmisc_1(u1_struct_0(X0),u1_struct_0(X0))) | ~v2_e_siec(X0) | ~l1_e_siec(X0) [cnf transformation 43] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
68. ~r1_tarski(X2,k2_zfmisc_1(X0,X1)) | r1_tarski(k9_xtuple_0(X2),X0) | ~v1_relat_1(X2) [cnf transformation 39] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
69. ~r1_tarski(X2,k2_zfmisc_1(X0,X1)) | r1_tarski(k10_xtuple_0(X2),X1) | ~v1_relat_1(X2) [cnf transformation 39] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. 1 <=> r1_tarski(k9_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0)) [avatar definition]
81. ~r1_tarski(k9_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0)) <- (~1) [avatar component clause 79] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. 2 <=> r1_tarski(k10_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0)) [avatar definition]
85. ~r1_tarski(k10_xtuple_0(u1_e_siec(sK0)),u1_struct_0(sK0)) <- (~2) [avatar component clause 83] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
87. 3 <=> r1_tarski(k9_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) [avatar definition]
89. ~r1_tarski(k9_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) <- (~3) [avatar component clause 87] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
91. 4 <=> r1_tarski(k10_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) [avatar definition]
93. ~r1_tarski(k10_xtuple_0(u2_e_siec(sK0)),u1_struct_0(sK0)) <- (~4) [avatar component clause 91] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
94. ~1 | ~2 | ~3 | ~4 [avatar split clause 55,91,87,83,79]
96. 5 <=> l1_e_siec(sK0) [avatar definition]
98. l1_e_siec(sK0) <- (5) [avatar component clause 96] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
99. 5 [avatar split clause 54,96]
101. 6 <=> v2_e_siec(sK0) [avatar definition]
103. v2_e_siec(sK0) <- (6) [avatar component clause 101] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
104. 6 [avatar split clause 53,101]
135. ~v2_e_siec(X0) | ~l1_e_siec(X0) | r1_tarski(k10_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)) | ~v1_relat_1(u1_e_siec(X0)) [resolution 58,69] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
136. ~v2_e_siec(X1) | ~l1_e_siec(X1) | r1_tarski(k9_xtuple_0(u1_e_siec(X1)),u1_struct_0(X1)) | ~v1_relat_1(u1_e_siec(X1)) [resolution 58,68] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
137. r1_tarski(k10_xtuple_0(u1_e_siec(X0)),u1_struct_0(X0)) | ~l1_e_siec(X0) | ~v2_e_siec(X0) [subsumption resolution 135,57] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
138. r1_tarski(k9_xtuple_0(u1_e_siec(X1)),u1_struct_0(X1)) | ~l1_e_siec(X1) | ~v2_e_siec(X1) [subsumption resolution 136,57] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
141. ~v2_e_siec(X0) | ~l1_e_siec(X0) | r1_tarski(k10_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) | ~v1_relat_1(u2_e_siec(X0)) [resolution 59,69] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
142. ~v2_e_siec(X1) | ~l1_e_siec(X1) | r1_tarski(k9_xtuple_0(u2_e_siec(X1)),u1_struct_0(X1)) | ~v1_relat_1(u2_e_siec(X1)) [resolution 59,68] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
143. r1_tarski(k10_xtuple_0(u2_e_siec(X0)),u1_struct_0(X0)) | ~l1_e_siec(X0) | ~v2_e_siec(X0) [subsumption resolution 141,56] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
144. r1_tarski(k9_xtuple_0(u2_e_siec(X1)),u1_struct_0(X1)) | ~l1_e_siec(X1) | ~v2_e_siec(X1) [subsumption resolution 142,56] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
161. ~l1_e_siec(sK0) | ~v2_e_siec(sK0) <- (~1) [resolution 138,81] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
164. ~v2_e_siec(sK0) <- (~1, 5) [subsumption resolution 161,98] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
165. $false <- (~1, 5, 6) [subsumption resolution 164,103] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
166. 1 | ~5 | ~6 [avatar contradiction clause 165]
178. ~l1_e_siec(sK0) | ~v2_e_siec(sK0) <- (~2) [resolution 85,137] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
180. ~v2_e_siec(sK0) <- (~2, 5) [subsumption resolution 178,98] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
181. $false <- (~2, 5, 6) [subsumption resolution 180,103] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
182. 2 | ~5 | ~6 [avatar contradiction clause 181]
186. ~l1_e_siec(sK0) | ~v2_e_siec(sK0) <- (~3) [resolution 144,89] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
189. ~v2_e_siec(sK0) <- (~3, 5) [subsumption resolution 186,98] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
190. $false <- (~3, 5, 6) [subsumption resolution 189,103] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
191. 3 | ~5 | ~6 [avatar contradiction clause 190]
194. ~l1_e_siec(sK0) | ~v2_e_siec(sK0) <- (~4) [resolution 93,143] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
196. ~v2_e_siec(sK0) <- (~4, 5) [subsumption resolution 194,98] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
197. $false <- (~4, 5, 6) [subsumption resolution 196,103] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
198. 4 | ~5 | ~6 [avatar contradiction clause 197]
199. $false [avatar sat refutation 94,99,104,166,182,191,198] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:79,thDist:-79}
% SZS output end Proof for e_siec__t14_e_siec
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% >>> Preprocessing
% Introduced skolems: 5
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 25
% Generated clauses: 98
% Active clauses: 43
% Passive clauses: 53
% Final active clauses: 40
% Final passive clauses: 10
% Main loop iterations started: 43
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 14
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 33
% Unit resulting resolution: 4
% Forward superposition: 1
% Backward superposition: 4
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 6
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 6
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

