i: [1,226,0,2,1,1,0,0] t1_mod_3
a: [3,331,0,2,1,226]
s: 331
i: [1,225,0,2,1,1,0,0] t1_mod_3
a: [3,336,0,2,1,225]
s: 336
i: [1,224,0,2,1,1,0,0] t1_mod_3
a: [3,341,0,2,1,224]
s: 341
i: [1,223,0,2,1,1,0,0] t1_mod_3
a: [3,346,0,2,1,223]
s: 346
i: [1,222,0,2,1,1,0,0] t1_mod_3
a: [3,351,0,2,1,222]
s: 351
i: [1,221,0,2,1,1,0,0] t1_mod_3
a: [3,356,0,2,1,221]
s: 356
i: [1,227,0,7,1,1,0,0] t1_mod_3
a: [3,326,0,7,1,227]
s: 326
i: [1,228,0,2,1,0,0,255] spc1_boole
a: [3,361,0,2,1,228]
s: 361
i: [1,229,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,366,0,2,1,229]
s: 366
i: [1,296,0,2,1,0,0,1] existence_l6_algstr_0
a: [3,376,0,2,1,296]
s: 376
i: [1,297,0,2,1,0,0,3] existence_l4_algstr_0
a: [3,381,0,2,1,297]
s: 381
i: [1,298,0,2,1,0,0,2] existence_l5_algstr_0
a: [3,386,0,2,1,298]
s: 386
i: [1,299,0,2,1,0,0,4] existence_l3_algstr_0
a: [3,391,0,2,1,299]
s: 391
i: [1,300,0,2,1,0,0,255] existence_l1_algstr_0
a: [3,396,0,2,1,300]
s: 396
i: [1,301,0,2,1,0,0,2] existence_l4_struct_0
a: [3,401,0,2,1,301]
s: 401
i: [1,302,0,2,1,0,0,2] existence_l2_algstr_0
a: [3,406,0,2,1,302]
s: 406
i: [1,303,0,2,1,0,0,255] rc4_ordinal1
a: [3,411,0,2,1,303]
s: 411
i: [1,304,0,2,1,0,0,2] existence_l3_struct_0
a: [3,416,0,2,1,304]
s: 416
i: [1,305,0,2,1,0,0,2] existence_l2_struct_0
a: [3,421,0,2,1,305]
s: 421
i: [1,306,0,2,1,0,0,3] existence_l1_struct_0
a: [3,426,0,2,1,306]
s: 426
i: [1,307,0,2,1,0,0,255] rc1_ordinal1
a: [3,431,0,2,1,307]
s: 431
i: [1,311,0,2,1,0,0,255] rc2_ordinal1
a: [3,436,0,2,1,311]
s: 436
i: [1,310,0,2,1,0,0,255] rc2_ordinal1
a: [3,441,0,2,1,310]
s: 441
i: [1,309,0,2,1,0,0,255] rc2_ordinal1
a: [3,446,0,2,1,309]
s: 446
i: [1,308,0,2,1,0,0,255] rc2_ordinal1
a: [3,451,0,2,1,308]
s: 451
i: [1,312,0,2,1,0,0,255] rc8_finseq_1
a: [3,456,0,2,1,312]
s: 456
i: [1,314,0,2,1,0,0,255] rc5_ordinal1
a: [3,461,0,2,1,314]
s: 461
i: [1,313,0,2,1,0,0,255] rc5_ordinal1
a: [3,466,0,2,1,313]
s: 466
i: [1,316,0,2,1,0,0,255] rc1_finset_1
a: [3,471,0,2,1,316]
s: 471
i: [1,315,0,2,1,0,0,255] rc1_finset_1
a: [3,476,0,2,1,315]
s: 476
i: [1,319,0,2,1,0,0,255] rc8_finset_1
a: [3,481,0,2,1,319]
s: 481
i: [1,318,0,2,1,0,0,255] rc8_finset_1
a: [3,486,0,2,1,318]
s: 486
i: [1,317,0,2,1,0,0,255] rc8_finset_1
a: [3,491,0,2,1,317]
s: 491
i: [1,322,0,2,1,0,0,2] rc20_struct_0
a: [3,496,0,2,1,322]
s: 496
i: [1,321,0,2,1,0,0,2] rc20_struct_0
a: [3,501,0,2,1,321]
s: 501
i: [1,320,0,2,1,0,0,2] rc20_struct_0
a: [3,506,0,2,1,320]
s: 506
i: [1,230,0,3,1,0,0,255] d2_xboole_0
a: [3,371,0,3,1,230]
s: 371
r: 366
d: [2,508,0,2,1,2,23,366,371]
a: [3,512,0,2,1,508]
s: 512
i: [1,231,0,4,2,0,0,255] cc5_finset_1
s: 231
i: [1,232,0,4,2,0,0,255] cc4_finset_1
s: 232
i: [1,233,0,4,2,0,0,4] dt_l3_algstr_0
s: 233
d: [2,514,1,2,1,1,40,233,391]
a: [3,518,1,2,1,514]
s: 518
i: [1,234,0,4,2,0,0,255] dt_l1_algstr_0
s: 234
d: [2,520,1,2,1,1,40,234,396]
a: [3,524,1,2,1,520]
s: 524
i: [1,235,0,4,2,0,0,2] dt_l5_algstr_0
s: 235
i: [1,236,0,4,2,0,0,2] dt_l5_algstr_0
s: 236
d: [2,526,1,2,1,1,40,235,386]
a: [3,530,1,2,1,526]
s: 530
i: [1,237,0,4,2,0,0,1] dt_l6_algstr_0
s: 237
d: [2,538,1,2,1,1,40,237,331]
a: [3,543,1,2,1,538]
s: 543
i: [1,238,0,4,2,0,0,1] dt_l6_algstr_0
s: 238
i: [1,239,0,4,2,0,0,3] dt_l4_algstr_0
s: 239
d: [2,550,1,2,1,1,40,238,331]
a: [3,555,1,2,1,550]
s: 555
i: [1,240,0,4,2,0,0,3] dt_l4_algstr_0
s: 240
d: [2,574,2,2,1,1,40,555,236]
a: [3,579,2,2,1,574]
s: 579
i: [1,242,0,4,2,0,0,2] dt_l4_struct_0
s: 242
i: [1,243,0,4,2,0,0,2] dt_l4_struct_0
s: 243
d: [2,575,2,2,1,1,40,555,235]
a: [3,584,2,2,1,575]
s: 584
i: [1,245,0,4,2,0,0,255] dt_l2_algstr_0
s: 245
d: [2,599,3,2,1,1,40,242,579]
a: [3,608,3,2,1,599]
s: 608
i: [1,246,0,4,2,0,0,255] dt_l2_algstr_0
s: 246
i: [1,247,0,4,2,0,0,255] cc6_ordinal1
s: 247
d: [2,639,2,2,1,1,40,246,543]
a: [3,648,2,2,1,639]
s: 648
i: [1,248,0,4,2,0,0,2] dt_l3_struct_0
s: 248
d: [2,611,3,2,1,1,40,243,579]
a: [3,620,3,2,1,611]
s: 620
i: [1,250,0,4,2,0,0,255] cc6_finset_1
s: 250
i: [1,251,0,4,2,0,0,255] cc7_ordinal1
s: 251
d: [2,622,3,2,1,1,40,584,240]
a: [3,627,3,2,1,622]
s: 627
i: [1,252,0,4,2,0,0,255] cc3_ordinal1
s: 252
d: [2,662,3,2,1,1,40,648,234]
a: [3,666,3,2,1,662]
s: 666
i: [1,253,0,4,2,0,0,255] cc1_finset_1
s: 253
i: [1,255,0,4,2,0,0,2] dt_l2_struct_0
s: 255
d: [2,532,1,2,1,1,40,236,386]
a: [3,536,1,2,1,532]
s: 536
i: [1,274,0,4,2,0,0,255] cc1_ordinal1
s: 274
d: [2,539,1,2,1,1,40,237,376]
a: [3,548,1,2,1,539]
s: 548
i: [1,275,0,4,2,0,0,255] cc1_ordinal1
s: 275
i: [1,293,0,4,1,0,0,2] existence_m1_subset_1
s: 293
d: [2,551,1,2,1,1,40,238,376]
a: [3,560,1,2,1,551]
s: 560
i: [1,254,0,5,2,0,0,255] t6_boole
s: 254
d: [2,562,1,2,1,1,40,239,381]
a: [3,567,1,2,1,562]
s: 567
i: [1,292,0,5,2,0,0,255] t7_boole
s: 292
i: [1,244,0,6,3,0,0,1] cc3_struct_0
s: 244
d: [2,770,3,4,2,1,40,244,579]
d: [2,781,3,2,1,2,21,770,351]
a: [3,785,3,2,1,781]
s: 785
i: [1,256,0,6,2,0,0,1] fc10_struct_0
s: 256
d: [2,586,1,2,1,1,40,240,381]
a: [3,591,1,2,1,586]
s: 591
i: [1,259,0,6,2,0,0,2] rc17_struct_0
s: 259
i: [1,260,0,6,3,0,0,255] cc7_struct_0
s: 260
d: [2,598,1,2,1,1,40,242,401]
a: [3,603,1,2,1,598]
s: 603
i: [1,261,0,6,3,0,0,255] cc5_struct_0
s: 261
d: [2,610,1,2,1,1,40,243,401]
a: [3,615,1,2,1,610]
s: 615
i: [1,263,0,6,3,0,0,2] cc17_struct_0
s: 263
i: [1,264,0,6,3,0,0,2] cc2_struct_0
s: 264
d: [2,630,1,2,1,1,40,245,406]
a: [3,635,1,2,1,630]
s: 635
i: [1,265,0,6,3,0,0,2] cc1_struct_0
s: 265
d: [2,638,1,2,1,1,40,246,406]
a: [3,643,1,2,1,638]
s: 643
i: [1,268,0,6,3,0,0,255] cc4_struct_0
s: 268
i: [1,269,0,6,3,0,0,255] cc6_struct_0
s: 269
d: [2,650,1,2,1,1,40,247,411]
a: [3,655,1,2,1,650]
s: 655
i: [1,285,0,6,3,0,0,255] cc2_ordinal1
s: 285
d: [2,651,1,2,1,1,40,247,461]
a: [3,660,1,2,1,651]
s: 660
i: [1,289,0,6,2,0,0,255] t1_subset
s: 289
i: [1,294,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 294
d: [2,668,1,2,1,1,40,248,416]
a: [3,672,1,2,1,668]
s: 672
i: [1,241,0,7,3,0,0,255] cc7_finset_1
s: 241
d: [2,676,1,2,1,1,40,250,512]
a: [3,680,1,2,1,676]
s: 680
i: [1,249,0,7,2,0,0,1] dt_k5_struct_0
s: 249
i: [1,257,0,7,2,0,0,1] dt_k4_struct_0
s: 257
d: [2,682,1,2,1,1,40,251,512]
a: [3,686,1,2,1,682]
s: 686
i: [1,258,0,7,2,0,0,2] rc17_struct_0
s: 258
d: [2,690,1,2,1,1,40,252,512]
a: [3,694,1,2,1,690]
s: 694
i: [1,266,0,7,3,0,0,255] cc8_struct_0
s: 266
i: [1,270,0,7,3,0,0,255] cc11_struct_0
s: 270
d: [2,696,1,2,1,1,40,253,512]
a: [3,700,1,2,1,696]
s: 700
i: [1,271,0,7,3,0,0,255] cc11_struct_0
s: 271
d: [2,702,1,2,1,1,40,255,421]
a: [3,708,1,2,1,702]
s: 708
i: [1,272,0,7,3,0,0,255] cc9_struct_0
s: 272
i: [1,276,0,7,3,0,0,255] cc5_ordinal1
s: 276
d: [2,703,1,2,1,1,40,255,506]
a: [3,713,1,2,1,703]
s: 713
i: [1,278,0,7,3,0,0,255] fc9_struct_0
s: 278
d: [2,723,1,2,1,1,40,274,431]
a: [3,728,1,2,1,723]
s: 728
i: [1,281,0,7,3,0,0,2] fc2_struct_0
s: 281
i: [1,284,0,7,3,0,0,255] fc6_struct_0
s: 284
d: [2,742,1,2,1,1,40,275,431]
a: [3,747,1,2,1,742]
s: 747
i: [1,286,0,7,3,0,0,255] fc8_struct_0
s: 286
d: [2,769,1,4,2,1,40,244,401]
a: [3,775,1,2,1,769]
s: 775
i: [1,287,0,7,3,0,0,2] fc1_struct_0
s: 287
i: [1,288,0,7,3,0,0,255] fc7_struct_0
s: 288
d: [2,798,1,4,2,1,40,260,496]
a: [3,802,1,2,1,798]
s: 802
i: [1,291,0,7,3,0,0,1] t8_boole
s: 291
d: [2,813,1,4,2,1,40,261,426]
a: [3,835,1,2,1,813]
s: 835
i: [1,277,0,8,3,0,0,1] fc11_struct_0
s: 277
i: [1,283,0,8,3,0,0,2] rc19_struct_0
s: 283
d: [2,563,2,2,1,1,40,239,530]
a: [3,572,2,2,1,563]
s: 572
i: [1,290,0,8,3,0,0,255] t2_subset
s: 290
d: [2,587,2,2,1,1,40,240,530]
a: [3,596,2,2,1,587]
s: 596
i: [1,273,0,9,4,0,0,255] cc10_struct_0
s: 273
i: [1,282,0,9,3,0,0,2] rc19_struct_0
s: 282
d: [2,716,2,2,1,1,40,536,242]
a: [3,721,2,2,1,716]
s: 721
i: [1,295,0,12,3,0,0,1] dt_k4_algstr_0
s: 295
d: [2,995,1,12,3,1,44,295,326]
d: [2,996,1,10,2,2,21,995,543]
a: [3,1000,1,5,1,996]
s: 1000
d: [2,1006,4,0,0,2,52,620,1000,249]
e: 1006
d: [2,1007,2,2,1,1,40,1000,249]
d: [2,1009,2,0,0,2,21,1007,620]
e: 1009
r: 1000
i: [1,279,0,17,6,0,0,1] t12_rlvect_1
s: 279
i: [1,280,0,21,7,0,0,1] t17_rlvect_1
s: 280
a: [3,999,1,5,1,996]
s: 999
a: [3,1004,1,5,1,996]
s: 1004
d: [2,1028,2,8,2,1,40,999,290]
a: [3,1038,2,3,1,1028]
s: 1038
d: [2,1076,4,0,0,3,52,356,666,1038,281]
e: 1076
d: [2,1077,3,4,2,1,40,1038,281]
d: [2,1083,3,2,1,2,21,1077,666]
d: [2,1084,3,0,0,3,21,1083,356]
e: 1084
r: 1038
d: [2,1013,1,22,7,1,44,280,326]
d: [2,1018,1,20,6,2,21,1013,543]
d: [2,1019,1,18,5,3,21,1018,336]
d: [2,1020,1,16,4,4,21,1019,341]
d: [2,1021,1,14,3,5,21,1020,346]
d: [2,1022,1,12,2,6,21,1021,356]
a: [3,1026,1,7,1,1022]
s: 1026
d: [2,911,1,5,2,0,40,241,293]
s: 911
a: [3,1037,2,3,1,1028]
s: 1037
d: [2,956,1,5,2,0,40,276,293]
s: 956
d: [2,1029,2,6,2,1,40,999,276]
a: [3,1043,2,3,1,1029]
s: 1043
d: [2,980,1,6,2,0,40,290,293]
s: 980
d: [2,761,1,7,3,0,44,254,254]
s: 761
d: [2,1030,2,6,2,1,40,999,241]
a: [3,1052,2,3,1,1030]
s: 1052
d: [2,963,1,7,3,0,40,287,253]
s: 963
a: [3,1034,2,5,1,1028]
s: 1034
d: [2,964,1,7,3,0,40,287,252]
s: 964
d: [2,965,1,7,3,0,40,287,251]
s: 965
d: [2,1058,2,8,2,1,40,1004,290]
a: [3,1064,2,5,1,1058]
s: 1064
d: [2,966,1,7,3,0,40,287,250]
s: 966
d: [2,1086,2,15,6,1,44,1026,279]
d: [2,1090,2,13,5,2,21,1086,543]
d: [2,1091,2,11,4,3,21,1090,336]
d: [2,1092,2,9,3,4,21,1091,341]
d: [2,1093,2,7,2,5,21,1092,346]
d: [2,1094,2,5,1,6,21,1093,356]
a: [3,1098,2,5,1,1094]
s: 1098
r: 1034
r: 1026
r: 999
r: 326
d: [2,1136,1,7,1,2,23,1026,1098]
a: [3,1144,1,7,1,1136]
s: 1144
d: [2,969,1,7,3,0,40,288,231]
s: 969
d: [2,1139,3,8,3,1,44,277,1098]
d: [2,1147,3,6,2,2,21,1139,579]
d: [2,1148,3,4,1,3,21,1147,351]
a: [3,1152,3,4,1,1148]
s: 1152
d: [2,1158,4,0,0,2,52,608,1152,256]
e: 1158
d: [2,1159,4,2,1,1,40,1152,256]
d: [2,1161,4,0,0,2,21,1159,608]
e: 1161
d: [2,623,3,2,1,1,40,584,239]
d: [2,631,2,2,1,1,40,245,543]
d: [2,1087,2,15,6,1,44,279,1026]
d: [2,1100,2,13,5,2,21,1087,543]
d: [2,1101,2,11,4,3,21,1100,336]
d: [2,1102,2,9,3,4,21,1101,341]
d: [2,1103,2,7,2,5,21,1102,346]
d: [2,1104,2,5,1,6,21,1103,356]
f: 1161,1148,1104,1022,1009,631,623,575,574,550,538,221,222,223,224,225,226,227
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mod_3__t1_mod_3
% SZS output start Proof for mod_3__t1_mod_3
1. ! [X0] : ((l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0)) => k4_struct_0(X0) != k4_algstr_0(X0,k5_struct_0(X0))) [input t1_mod_3]
2. ~! [X0] : ((l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0)) => k4_struct_0(X0) != k4_algstr_0(X0,k5_struct_0(X0))) [negated conjecture 1]
36. ! [X0] : (l3_struct_0(X0) => m1_subset_1(k5_struct_0(X0),u1_struct_0(X0))) [input dt_k5_struct_0]
39. ! [X0] : (l2_algstr_0(X0) => (l1_algstr_0(X0) & l2_struct_0(X0))) [input dt_l2_algstr_0]
43. ! [X0] : (l4_algstr_0(X0) => (l3_algstr_0(X0) & l3_struct_0(X0))) [input dt_l4_algstr_0]
45. ! [X0] : (l5_algstr_0(X0) => (l4_struct_0(X0) & l4_algstr_0(X0))) [input dt_l5_algstr_0]
46. ! [X0] : (l6_algstr_0(X0) => (l5_algstr_0(X0) & l2_algstr_0(X0))) [input dt_l6_algstr_0]
61. ! [X0] : (l2_struct_0(X0) => v9_struct_0(k4_struct_0(X0),X0)) [input fc10_struct_0]
62. ! [X0] : ((l4_struct_0(X0) & ~v6_struct_0(X0)) => ~v9_struct_0(k5_struct_0(X0),X0)) [input fc11_struct_0]
80. ! [X0] : ((l2_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v13_algstr_0(X0) & ~v2_struct_0(X0)) => k4_struct_0(X0) = k4_algstr_0(X0,k4_struct_0(X0))) [input t12_rlvect_1]
81. ! [X0] : ((l2_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v13_algstr_0(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => k4_algstr_0(X0,k4_algstr_0(X0,X1)) = X1)) [input t17_rlvect_1]
97. ? [X0] : (k4_struct_0(X0) = k4_algstr_0(X0,k5_struct_0(X0)) & (l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
98. ? [X0] : (k4_struct_0(X0) = k4_algstr_0(X0,k5_struct_0(X0)) & l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0)) [flattening 97]
103. ! [X0] : ((l4_struct_0(X0) & l4_algstr_0(X0)) | ~l5_algstr_0(X0)) [ennf transformation 45]
104. ! [X0] : ((l5_algstr_0(X0) & l2_algstr_0(X0)) | ~l6_algstr_0(X0)) [ennf transformation 46]
105. ! [X0] : ((l3_algstr_0(X0) & l3_struct_0(X0)) | ~l4_algstr_0(X0)) [ennf transformation 43]
110. ! [X0] : ((l1_algstr_0(X0) & l2_struct_0(X0)) | ~l2_algstr_0(X0)) [ennf transformation 39]
113. ! [X0] : (m1_subset_1(k5_struct_0(X0),u1_struct_0(X0)) | ~l3_struct_0(X0)) [ennf transformation 36]
120. ! [X0] : (v9_struct_0(k4_struct_0(X0),X0) | ~l2_struct_0(X0)) [ennf transformation 61]
147. ! [X0] : (~v9_struct_0(k5_struct_0(X0),X0) | (~l4_struct_0(X0) | v6_struct_0(X0))) [ennf transformation 62]
148. ! [X0] : (~v9_struct_0(k5_struct_0(X0),X0) | ~l4_struct_0(X0) | v6_struct_0(X0)) [flattening 147]
151. ! [X0] : (k4_struct_0(X0) = k4_algstr_0(X0,k4_struct_0(X0)) | (~l2_algstr_0(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 80]
152. ! [X0] : (k4_struct_0(X0) = k4_algstr_0(X0,k4_struct_0(X0)) | ~l2_algstr_0(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0)) [flattening 151]
153. ! [X0] : (! [X1] : (k4_algstr_0(X0,k4_algstr_0(X0,X1)) = X1 | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_algstr_0(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0))) [ennf transformation 81]
154. ! [X0] : (! [X1] : (k4_algstr_0(X0,k4_algstr_0(X0,X1)) = X1 | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_algstr_0(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0)) [flattening 153]
177. ? [X0] : (k4_struct_0(X0) = k4_algstr_0(X0,k5_struct_0(X0)) & l6_algstr_0(X0) & v4_rlvect_1(X0) & v3_rlvect_1(X0) & v13_algstr_0(X0) & ~v6_struct_0(X0) & ~v2_struct_0(X0)) => (k4_struct_0(sK0) = k4_algstr_0(sK0,k5_struct_0(sK0)) & l6_algstr_0(sK0) & v4_rlvect_1(sK0) & v3_rlvect_1(sK0) & v13_algstr_0(sK0) & ~v6_struct_0(sK0) & ~v2_struct_0(sK0)) [choice axiom]
178. k4_struct_0(sK0) = k4_algstr_0(sK0,k5_struct_0(sK0)) & l6_algstr_0(sK0) & v4_rlvect_1(sK0) & v3_rlvect_1(sK0) & v13_algstr_0(sK0) & ~v6_struct_0(sK0) & ~v2_struct_0(sK0) [skolemisation 98,177]
221. ~v2_struct_0(sK0) [cnf transformation 178] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
222. ~v6_struct_0(sK0) [cnf transformation 178] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
223. v13_algstr_0(sK0) [cnf transformation 178] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
224. v3_rlvect_1(sK0) [cnf transformation 178] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
225. v4_rlvect_1(sK0) [cnf transformation 178] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
226. l6_algstr_0(sK0) [cnf transformation 178] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
227. k4_struct_0(sK0) = k4_algstr_0(sK0,k5_struct_0(sK0)) [cnf transformation 178] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
235. ~l5_algstr_0(X0) | l4_algstr_0(X0) [cnf transformation 103] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
236. ~l5_algstr_0(X0) | l4_struct_0(X0) [cnf transformation 103] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
237. ~l6_algstr_0(X0) | l2_algstr_0(X0) [cnf transformation 104] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
238. ~l6_algstr_0(X0) | l5_algstr_0(X0) [cnf transformation 104] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
239. ~l4_algstr_0(X0) | l3_struct_0(X0) [cnf transformation 105] {a:0,w:4,wCS:1600,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
245. ~l2_algstr_0(X0) | l2_struct_0(X0) [cnf transformation 110] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
249. m1_subset_1(k5_struct_0(X0),u1_struct_0(X0)) | ~l3_struct_0(X0) [cnf transformation 113] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
256. v9_struct_0(k4_struct_0(X0),X0) | ~l2_struct_0(X0) [cnf transformation 120] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
277. ~v9_struct_0(k5_struct_0(X0),X0) | ~l4_struct_0(X0) | v6_struct_0(X0) [cnf transformation 148] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
279. k4_struct_0(X0) = k4_algstr_0(X0,k4_struct_0(X0)) | ~l2_algstr_0(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 152] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
280. k4_algstr_0(X0,k4_algstr_0(X0,X1)) = X1 | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_algstr_0(X0) | ~v4_rlvect_1(X0) | ~v3_rlvect_1(X0) | ~v13_algstr_0(X0) | v2_struct_0(X0) [cnf transformation 154] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
324. 1 <=> k4_struct_0(sK0) = k4_algstr_0(sK0,k5_struct_0(sK0)) [avatar definition]
326. k4_struct_0(sK0) = k4_algstr_0(sK0,k5_struct_0(sK0)) <- (1) [avatar component clause 324] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
327. 1 [avatar split clause 227,324]
329. 2 <=> l6_algstr_0(sK0) [avatar definition]
331. l6_algstr_0(sK0) <- (2) [avatar component clause 329] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
332. 2 [avatar split clause 226,329]
334. 3 <=> v4_rlvect_1(sK0) [avatar definition]
336. v4_rlvect_1(sK0) <- (3) [avatar component clause 334] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
337. 3 [avatar split clause 225,334]
339. 4 <=> v3_rlvect_1(sK0) [avatar definition]
341. v3_rlvect_1(sK0) <- (4) [avatar component clause 339] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
342. 4 [avatar split clause 224,339]
344. 5 <=> v13_algstr_0(sK0) [avatar definition]
346. v13_algstr_0(sK0) <- (5) [avatar component clause 344] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
347. 5 [avatar split clause 223,344]
349. 6 <=> v6_struct_0(sK0) [avatar definition]
351. ~v6_struct_0(sK0) <- (~6) [avatar component clause 349] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
352. ~6 [avatar split clause 222,349]
354. 7 <=> v2_struct_0(sK0) [avatar definition]
356. ~v2_struct_0(sK0) <- (~7) [avatar component clause 354] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
357. ~7 [avatar split clause 221,354]
538. l2_algstr_0(sK0) <- (2) [resolution 237,331] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
541. 43 <=> l2_algstr_0(sK0) [avatar definition]
543. l2_algstr_0(sK0) <- (43) [avatar component clause 541] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
544. 43 | ~2 [avatar split clause 538,329,541]
550. l5_algstr_0(sK0) <- (2) [resolution 238,331] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
553. 45 <=> l5_algstr_0(sK0) [avatar definition]
555. l5_algstr_0(sK0) <- (45) [avatar component clause 553] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
556. 45 | ~2 [avatar split clause 550,329,553]
574. l4_struct_0(sK0) <- (45) [resolution 555,236] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
575. l4_algstr_0(sK0) <- (45) [resolution 555,235] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
577. 49 <=> l4_struct_0(sK0) [avatar definition]
579. l4_struct_0(sK0) <- (49) [avatar component clause 577] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
580. 49 | ~45 [avatar split clause 574,553,577]
582. 50 <=> l4_algstr_0(sK0) [avatar definition]
584. l4_algstr_0(sK0) <- (50) [avatar component clause 582] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
585. 50 | ~45 [avatar split clause 575,553,582]
606. 54 <=> l2_struct_0(sK0) [avatar definition]
608. l2_struct_0(sK0) <- (54) [avatar component clause 606] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
618. 56 <=> l3_struct_0(sK0) [avatar definition]
620. l3_struct_0(sK0) <- (56) [avatar component clause 618] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
623. l3_struct_0(sK0) <- (50) [resolution 584,239] {a:3,w:2,wCS:200,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
629. 56 | ~50 [avatar split clause 623,582,618]
631. l2_struct_0(sK0) <- (43) [resolution 245,543] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
637. 54 | ~43 [avatar split clause 631,541,606]
998. 109 <=> m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) [avatar definition]
1000. ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) <- (~109) [avatar component clause 998] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1007. ~l3_struct_0(sK0) <- (~109) [resolution 1000,249] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1009. $false <- (56, ~109) [subsumption resolution 1007,620] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1010. ~56 | 109 [avatar contradiction clause 1009]
1013. k5_struct_0(sK0) = k4_algstr_0(sK0,k4_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~l2_algstr_0(sK0) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (1) [superposition 280,326] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1018. k5_struct_0(sK0) = k4_algstr_0(sK0,k4_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (1, 43) [subsumption resolution 1013,543] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1019. k5_struct_0(sK0) = k4_algstr_0(sK0,k4_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (1, 3, 43) [subsumption resolution 1018,336] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1020. k5_struct_0(sK0) = k4_algstr_0(sK0,k4_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (1, 3, 4, 43) [subsumption resolution 1019,341] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1021. k5_struct_0(sK0) = k4_algstr_0(sK0,k4_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) | v2_struct_0(sK0) <- (1, 3, 4, 5, 43) [subsumption resolution 1020,346] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1022. k5_struct_0(sK0) = k4_algstr_0(sK0,k4_struct_0(sK0)) | ~m1_subset_1(k5_struct_0(sK0),u1_struct_0(sK0)) <- (1, 3, 4, 5, ~7, 43) [subsumption resolution 1021,356] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1024. 111 <=> k5_struct_0(sK0) = k4_algstr_0(sK0,k4_struct_0(sK0)) [avatar definition]
1026. k5_struct_0(sK0) = k4_algstr_0(sK0,k4_struct_0(sK0)) <- (111) [avatar component clause 1024] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1027. ~109 | 111 | ~1 | ~3 | ~4 | ~5 | 7 | ~43 [avatar split clause 1022,541,354,344,339,334,324,1024,998]
1087. k4_struct_0(sK0) = k5_struct_0(sK0) | ~l2_algstr_0(sK0) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (111) [superposition 279,1026] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1096. 121 <=> k4_struct_0(sK0) = k5_struct_0(sK0) [avatar definition]
1098. k4_struct_0(sK0) = k5_struct_0(sK0) <- (121) [avatar component clause 1096] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1100. k4_struct_0(sK0) = k5_struct_0(sK0) | ~v4_rlvect_1(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (43, 111) [subsumption resolution 1087,543] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1101. k4_struct_0(sK0) = k5_struct_0(sK0) | ~v3_rlvect_1(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (3, 43, 111) [subsumption resolution 1100,336] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1102. k4_struct_0(sK0) = k5_struct_0(sK0) | ~v13_algstr_0(sK0) | v2_struct_0(sK0) <- (3, 4, 43, 111) [subsumption resolution 1101,341] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1103. k4_struct_0(sK0) = k5_struct_0(sK0) | v2_struct_0(sK0) <- (3, 4, 5, 43, 111) [subsumption resolution 1102,346] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1104. k4_struct_0(sK0) = k5_struct_0(sK0) <- (3, 4, 5, ~7, 43, 111) [subsumption resolution 1103,356] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1105. 121 | ~3 | ~4 | ~5 | 7 | ~43 | ~111 [avatar split clause 1104,1024,541,354,344,339,334,1096]
1139. ~v9_struct_0(k4_struct_0(sK0),sK0) | ~l4_struct_0(sK0) | v6_struct_0(sK0) <- (121) [superposition 277,1098] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1147. ~v9_struct_0(k4_struct_0(sK0),sK0) | v6_struct_0(sK0) <- (49, 121) [subsumption resolution 1139,579] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1148. ~v9_struct_0(k4_struct_0(sK0),sK0) <- (~6, 49, 121) [subsumption resolution 1147,351] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1150. 125 <=> v9_struct_0(k4_struct_0(sK0),sK0) [avatar definition]
1152. ~v9_struct_0(k4_struct_0(sK0),sK0) <- (~125) [avatar component clause 1150] {a:3,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1153. ~125 | 6 | ~49 | ~121 [avatar split clause 1148,1096,577,349,1150]
1159. ~l2_struct_0(sK0) <- (~125) [resolution 1152,256] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1161. $false <- (54, ~125) [subsumption resolution 1159,608] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1162. ~54 | 125 [avatar contradiction clause 1161]
1163. $false [avatar sat refutation 327,332,337,342,347,352,357,544,556,580,585,629,637,1010,1027,1105,1153,1162] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:710,thDist:-710}
% SZS output end Proof for mod_3__t1_mod_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 85
% >>> Preprocessing
% Introduced skolems: 22
% Pure predicates: 5
% 
% >>> Saturation
% Initial clauses: 102
% Generated clauses: 420
% Active clauses: 156
% Passive clauses: 205
% Final active clauses: 149
% Final passive clauses: 43
% Main loop iterations started: 162
% 
% >>> Simplifying Inferences
% Duplicate literals: 15
% Fw subsumption resolutions: 23
% Bw demodulations: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 37
% 
% >>> Generating Inferences
% Binary resolution: 144
% Unit resulting resolution: 3
% Forward superposition: 10
% Backward superposition: 10
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 21
% Split components: 42
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 135
% SAT solver unit clauses: 37
% SAT solver binary clauses: 66
% 
% Memory used [KB]: 10106
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

