i: [1,59,0,2,1,1,0,0] t24_pboole
a: [3,97,0,2,1,59]
s: 97
i: [1,57,0,2,1,1,0,0] t24_pboole
a: [3,107,0,2,1,57]
s: 107
i: [1,55,0,2,1,1,0,0] t24_pboole
a: [3,117,0,2,1,55]
s: 117
i: [1,53,0,2,1,1,0,0] t24_pboole
a: [3,127,0,2,1,53]
s: 127
i: [1,51,0,2,1,1,0,0] t24_pboole
a: [3,137,0,2,1,51]
s: 137
i: [1,49,0,2,1,1,0,0] t24_pboole
a: [3,147,0,2,1,49]
s: 147
i: [1,60,0,3,1,1,0,0] t24_pboole
a: [3,92,0,3,1,60]
s: 92
i: [1,58,0,3,1,1,0,0] t24_pboole
a: [3,102,0,3,1,58]
s: 102
i: [1,56,0,3,1,1,0,0] t24_pboole
a: [3,112,0,3,1,56]
s: 112
i: [1,54,0,3,1,1,0,0] t24_pboole
a: [3,122,0,3,1,54]
s: 122
i: [1,52,0,3,1,1,0,0] t24_pboole
a: [3,132,0,3,1,52]
s: 132
i: [1,50,0,3,1,1,0,0] t24_pboole
a: [3,142,0,3,1,50]
s: 142
i: [1,83,0,2,1,0,0,1] rc1_funct_1
a: [3,152,0,2,1,83]
s: 152
i: [1,82,0,2,1,0,0,1] rc1_funct_1
a: [3,157,0,2,1,82]
s: 157
i: [1,61,0,10,1,1,0,0] t24_pboole
a: [3,87,0,10,1,61]
s: 87
i: [1,65,0,3,1,0,0,1] rc1_pboole
s: 65
i: [1,67,0,3,1,0,0,1] rc1_pboole
s: 67
i: [1,66,0,4,1,0,0,1] rc1_pboole
s: 66
i: [1,68,0,4,1,0,0,1] rc1_pboole
s: 68
i: [1,69,0,24,9,0,0,1] reflexivity_r2_pboole
s: 69
i: [1,74,0,25,9,0,0,1] dt_k3_pboole
s: 74
i: [1,76,0,25,9,0,0,1] dt_k3_pboole
s: 76
i: [1,78,0,25,9,0,0,1] dt_k2_pboole
s: 78
i: [1,80,0,25,9,0,0,1] dt_k2_pboole
s: 80
i: [1,70,0,26,9,0,0,255] idempotence_k2_pboole
s: 70
i: [1,71,0,26,9,0,0,255] idempotence_k3_pboole
s: 71
i: [1,75,0,26,9,0,0,1] dt_k3_pboole
s: 75
i: [1,77,0,26,9,0,0,1] dt_k3_pboole
s: 77
i: [1,79,0,26,9,0,0,1] dt_k2_pboole
s: 79
i: [1,81,0,26,9,0,0,1] dt_k2_pboole
s: 81
i: [1,62,0,27,9,0,0,1] t14_pboole
s: 62
i: [1,63,0,27,9,0,0,1] t15_pboole
s: 63
i: [1,72,0,29,9,0,0,255] commutativity_k2_pboole
s: 72
d: [2,184,1,47,17,0,44,62,72]
d: [2,193,1,27,9,0,18,184]
s: 193
i: [1,73,0,29,9,0,0,255] commutativity_k3_pboole
s: 73
i: [1,64,0,42,15,0,0,1] t13_pboole
s: 64
d: [2,179,1,47,17,0,44,62,72]
d: [2,198,1,27,9,0,18,179]
s: 198
d: [2,210,1,47,17,0,44,63,73]
d: [2,219,1,27,9,0,18,210]
s: 219
d: [2,205,1,47,17,0,44,63,73]
d: [2,224,1,27,9,0,18,205]
s: 224
d: [2,226,1,73,22,0,40,64,62]
d: [2,231,1,63,18,0,18,226]
d: [2,233,1,57,17,0,21,231,81]
d: [2,234,1,52,16,0,21,233,80]
d: [2,235,1,46,15,0,21,234,79]
d: [2,236,1,41,14,0,21,235,78]
s: 236
d: [2,228,1,73,22,0,40,64,63]
d: [2,229,1,63,18,0,18,228]
d: [2,241,1,57,17,0,21,229,77]
d: [2,242,1,52,16,0,21,241,76]
d: [2,243,1,46,15,0,21,242,75]
d: [2,244,1,41,14,0,21,243,74]
s: 244
d: [2,257,2,49,13,1,40,236,87]
d: [2,266,2,46,12,2,21,257,132]
d: [2,267,2,44,11,3,21,266,137]
d: [2,268,2,41,10,4,21,267,142]
d: [2,269,2,39,9,5,21,268,147]
d: [2,270,2,36,8,6,21,269,92]
d: [2,271,2,34,7,7,21,270,97]
d: [2,272,2,31,6,8,21,271,102]
d: [2,273,2,29,5,9,21,272,107]
a: [3,277,2,5,1,273]
s: 277
d: [2,325,3,0,0,9,52,147,127,137,117,142,132,122,112,277,74]
e: 325
d: [2,326,3,20,8,1,40,277,74]
d: [2,328,3,17,7,2,21,326,112]
d: [2,329,3,15,6,3,21,328,117]
d: [2,330,3,12,5,4,21,329,122]
d: [2,331,3,10,4,5,21,330,127]
d: [2,332,3,7,3,6,21,331,132]
d: [2,333,3,5,2,7,21,332,137]
d: [2,334,3,2,1,8,21,333,142]
d: [2,335,3,0,0,9,21,334,147]
e: 335
r: 277
a: [3,276,2,5,1,273]
s: 276
a: [3,281,2,6,1,273]
s: 281
d: [2,337,3,0,0,9,52,147,127,137,117,142,132,122,112,281,75]
e: 337
d: [2,338,3,20,8,1,40,281,75]
d: [2,340,3,17,7,2,21,338,112]
d: [2,341,3,15,6,3,21,340,117]
d: [2,342,3,12,5,4,21,341,122]
d: [2,343,3,10,4,5,21,342,127]
d: [2,344,3,7,3,6,21,343,132]
d: [2,345,3,5,2,7,21,344,137]
d: [2,346,3,2,1,8,21,345,142]
d: [2,347,3,0,0,9,21,346,147]
e: 347
r: 281
a: [3,285,2,5,1,273]
s: 285
d: [2,349,3,0,0,9,52,147,127,137,117,142,132,122,112,285,76]
e: 349
d: [2,350,3,20,8,1,40,285,76]
d: [2,352,3,17,7,2,21,350,112]
d: [2,353,3,15,6,3,21,352,117]
d: [2,354,3,12,5,4,21,353,122]
d: [2,355,3,10,4,5,21,354,127]
d: [2,356,3,7,3,6,21,355,132]
d: [2,357,3,5,2,7,21,356,137]
d: [2,358,3,2,1,8,21,357,142]
d: [2,359,3,0,0,9,21,358,147]
e: 359
r: 285
a: [3,284,2,5,1,273]
s: 284
a: [3,280,2,6,1,273]
s: 280
a: [3,289,2,6,1,273]
s: 289
d: [2,361,3,0,0,9,52,147,127,137,117,142,132,122,112,289,77]
e: 361
d: [2,362,3,20,8,1,40,289,77]
d: [2,364,3,17,7,2,21,362,112]
d: [2,365,3,15,6,3,21,364,117]
d: [2,366,3,12,5,4,21,365,122]
d: [2,367,3,10,4,5,21,366,127]
d: [2,368,3,7,3,6,21,367,132]
d: [2,369,3,5,2,7,21,368,137]
d: [2,370,3,2,1,8,21,369,142]
d: [2,371,3,0,0,9,21,370,147]
e: 371
r: 289
a: [3,288,2,6,1,273]
s: 288
a: [3,293,2,7,1,273]
s: 293
d: [2,373,3,0,0,9,52,147,127,137,117,142,132,122,112,293,63]
e: 373
d: [2,374,3,20,8,1,40,293,63]
d: [2,376,3,17,7,2,21,374,112]
d: [2,377,3,15,6,3,21,376,117]
d: [2,378,3,12,5,4,21,377,122]
d: [2,379,3,10,4,5,21,378,127]
d: [2,380,3,7,3,6,21,379,132]
d: [2,381,3,5,2,7,21,380,137]
d: [2,382,3,2,1,8,21,381,142]
d: [2,383,3,0,0,9,21,382,147]
e: 383
f: 383,371,359,347,335,273,49,50,51,52,53,54,55,56,57,58,59,60,61
% Refutation found. Thanks to Tanya!
% SZS status Theorem for pboole__t24_pboole
% SZS output start Proof for pboole__t24_pboole
1. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => r2_pboole(X0,k3_pboole(X0,X1,X2),k2_pboole(X0,X1,X3))))) [input t24_pboole]
2. ~! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => r2_pboole(X0,k3_pboole(X0,X1,X2),k2_pboole(X0,X1,X3))))) [negated conjecture 1]
5. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2)))) [input dt_k2_pboole]
6. ! [X0,X1,X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2)))) [input dt_k3_pboole]
12. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,X3))))) [input t13_pboole]
13. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,X1,k2_pboole(X0,X1,X2)))) [input t14_pboole]
14. ! [X0] : ! [X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k3_pboole(X0,X1,X2),X1))) [input t15_pboole]
15. ~! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => r2_pboole(X0,k3_pboole(X0,X1,X2),k2_pboole(X0,X1,X3))))) [flattening 2]
16. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,X1,k2_pboole(X0,X1,X2)))) [flattening 13]
17. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => r2_pboole(X0,k3_pboole(X0,X1,X2),X1))) [flattening 14]
18. ! [X0,X1] : ((v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : ((v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) => ! [X3] : ((v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) => ((r2_pboole(X0,X2,X3) & r2_pboole(X0,X1,X2)) => r2_pboole(X0,X1,X3))))) [flattening 12]
19. ? [X0,X1] : (? [X2] : (? [X3] : (~r2_pboole(X0,k3_pboole(X0,X1,X2),k2_pboole(X0,X1,X3)) & (v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3))) & (v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2))) & (v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 15]
20. ? [X0,X1] : (? [X2] : (? [X3] : (~r2_pboole(X0,k3_pboole(X0,X1,X2),k2_pboole(X0,X1,X3)) & v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 19]
21. ! [X0,X1] : (! [X2] : (r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 16]
22. ! [X0,X1] : (! [X2] : (r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 21]
23. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 17]
24. ! [X0,X1] : (! [X2] : (r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 23]
25. ! [X0,X1] : (! [X2] : (! [X3] : ((r2_pboole(X0,X1,X3) | (~r2_pboole(X0,X2,X3) | ~r2_pboole(X0,X1,X2))) | (~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2))) | (~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 18]
26. ! [X0,X1] : (! [X2] : (! [X3] : (r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X2,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2)) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 25]
37. ! [X0,X1,X2] : ((v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 6]
38. ! [X0,X1,X2] : ((v1_partfun1(k3_pboole(X0,X1,X2),X0) & v1_funct_1(k3_pboole(X0,X1,X2)) & v4_relat_1(k3_pboole(X0,X1,X2),X0) & v1_relat_1(k3_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 37]
39. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | (~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 5]
40. ! [X0,X1,X2] : ((v1_partfun1(k2_pboole(X0,X1,X2),X0) & v1_funct_1(k2_pboole(X0,X1,X2)) & v4_relat_1(k2_pboole(X0,X1,X2),X0) & v1_relat_1(k2_pboole(X0,X1,X2))) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 39]
41. ? [X0,X1] : (? [X2] : (? [X3] : (~r2_pboole(X0,k3_pboole(X0,X1,X2),k2_pboole(X0,X1,X3)) & v1_partfun1(X3,X0) & v1_funct_1(X3) & v4_relat_1(X3,X0) & v1_relat_1(X3)) & v1_partfun1(X2,X0) & v1_funct_1(X2) & v4_relat_1(X2,X0) & v1_relat_1(X2)) & v1_partfun1(X1,X0) & v1_funct_1(X1) & v4_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (~r2_pboole(sK0,k3_pboole(sK0,sK1,X2),k2_pboole(sK0,sK1,X3)) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
42. ? [X2] : (? [X3] : (~r2_pboole(sK0,k3_pboole(sK0,sK1,X2),k2_pboole(sK0,sK1,X3)) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) & v1_partfun1(X2,sK0) & v1_funct_1(X2) & v4_relat_1(X2,sK0) & v1_relat_1(X2)) => (? [X3] : (~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),k2_pboole(sK0,sK1,X3)) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
43. ? [X3] : (~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),k2_pboole(sK0,sK1,X3)) & v1_partfun1(X3,sK0) & v1_funct_1(X3) & v4_relat_1(X3,sK0) & v1_relat_1(X3)) => (~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),k2_pboole(sK0,sK1,sK3)) & v1_partfun1(sK3,sK0) & v1_funct_1(sK3) & v4_relat_1(sK3,sK0) & v1_relat_1(sK3)) [choice axiom]
44. ((~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),k2_pboole(sK0,sK1,sK3)) & v1_partfun1(sK3,sK0) & v1_funct_1(sK3) & v4_relat_1(sK3,sK0) & v1_relat_1(sK3)) & v1_partfun1(sK2,sK0) & v1_funct_1(sK2) & v4_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_partfun1(sK1,sK0) & v1_funct_1(sK1) & v4_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 20,43,42,41]
49. v1_relat_1(sK1) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
50. v4_relat_1(sK1,sK0) [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
51. v1_funct_1(sK1) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
52. v1_partfun1(sK1,sK0) [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
53. v1_relat_1(sK2) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
54. v4_relat_1(sK2,sK0) [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
55. v1_funct_1(sK2) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
56. v1_partfun1(sK2,sK0) [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
57. v1_relat_1(sK3) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
58. v4_relat_1(sK3,sK0) [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
59. v1_funct_1(sK3) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
60. v1_partfun1(sK3,sK0) [cnf transformation 44] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
61. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),k2_pboole(sK0,sK1,sK3)) [cnf transformation 44] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
62. r2_pboole(X0,X1,k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 22] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. r2_pboole(X0,k3_pboole(X0,X1,X2),X1) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 24] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. ~r2_pboole(X0,X2,X3) | r2_pboole(X0,X1,X3) | ~r2_pboole(X0,X1,X2) | ~v1_partfun1(X3,X0) | ~v1_funct_1(X3) | ~v4_relat_1(X3,X0) | ~v1_relat_1(X3) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 26] {a:0,w:42,wCS:16800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
74. v1_relat_1(k3_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 38] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
75. v4_relat_1(k3_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 38] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
76. v1_funct_1(k3_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 38] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
77. v1_partfun1(k3_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 38] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
78. v1_relat_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 40] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. v4_relat_1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 40] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
80. v1_funct_1(k2_pboole(X0,X1,X2)) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 40] {a:0,w:25,wCS:10000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
81. v1_partfun1(k2_pboole(X0,X1,X2),X0) | ~v1_partfun1(X2,X0) | ~v1_funct_1(X2) | ~v4_relat_1(X2,X0) | ~v1_relat_1(X2) | ~v1_partfun1(X1,X0) | ~v1_funct_1(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 40] {a:0,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
85. 1 <=> r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),k2_pboole(sK0,sK1,sK3)) [avatar definition]
87. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),k2_pboole(sK0,sK1,sK3)) <- (~1) [avatar component clause 85] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
88. ~1 [avatar split clause 61,85]
90. 2 <=> v1_partfun1(sK3,sK0) [avatar definition]
92. v1_partfun1(sK3,sK0) <- (2) [avatar component clause 90] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
93. 2 [avatar split clause 60,90]
95. 3 <=> v1_funct_1(sK3) [avatar definition]
97. v1_funct_1(sK3) <- (3) [avatar component clause 95] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
98. 3 [avatar split clause 59,95]
100. 4 <=> v4_relat_1(sK3,sK0) [avatar definition]
102. v4_relat_1(sK3,sK0) <- (4) [avatar component clause 100] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
103. 4 [avatar split clause 58,100]
105. 5 <=> v1_relat_1(sK3) [avatar definition]
107. v1_relat_1(sK3) <- (5) [avatar component clause 105] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
108. 5 [avatar split clause 57,105]
110. 6 <=> v1_partfun1(sK2,sK0) [avatar definition]
112. v1_partfun1(sK2,sK0) <- (6) [avatar component clause 110] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
113. 6 [avatar split clause 56,110]
115. 7 <=> v1_funct_1(sK2) [avatar definition]
117. v1_funct_1(sK2) <- (7) [avatar component clause 115] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
118. 7 [avatar split clause 55,115]
120. 8 <=> v4_relat_1(sK2,sK0) [avatar definition]
122. v4_relat_1(sK2,sK0) <- (8) [avatar component clause 120] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
123. 8 [avatar split clause 54,120]
125. 9 <=> v1_relat_1(sK2) [avatar definition]
127. v1_relat_1(sK2) <- (9) [avatar component clause 125] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
128. 9 [avatar split clause 53,125]
130. 10 <=> v1_partfun1(sK1,sK0) [avatar definition]
132. v1_partfun1(sK1,sK0) <- (10) [avatar component clause 130] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
133. 10 [avatar split clause 52,130]
135. 11 <=> v1_funct_1(sK1) [avatar definition]
137. v1_funct_1(sK1) <- (11) [avatar component clause 135] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
138. 11 [avatar split clause 51,135]
140. 12 <=> v4_relat_1(sK1,sK0) [avatar definition]
142. v4_relat_1(sK1,sK0) <- (12) [avatar component clause 140] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
143. 12 [avatar split clause 50,140]
145. 13 <=> v1_relat_1(sK1) [avatar definition]
147. v1_relat_1(sK1) <- (13) [avatar component clause 145] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
148. 13 [avatar split clause 49,145]
226. r2_pboole(X4,X5,k2_pboole(X4,X6,X7)) | ~r2_pboole(X4,X5,X6) | ~v1_partfun1(k2_pboole(X4,X6,X7),X4) | ~v1_funct_1(k2_pboole(X4,X6,X7)) | ~v4_relat_1(k2_pboole(X4,X6,X7),X4) | ~v1_relat_1(k2_pboole(X4,X6,X7)) | ~v1_partfun1(X6,X4) | ~v1_funct_1(X6) | ~v4_relat_1(X6,X4) | ~v1_relat_1(X6) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~v1_partfun1(X7,X4) | ~v1_funct_1(X7) | ~v4_relat_1(X7,X4) | ~v1_relat_1(X7) | ~v1_partfun1(X6,X4) | ~v1_funct_1(X6) | ~v4_relat_1(X6,X4) | ~v1_relat_1(X6) [resolution 64,62] {a:1,w:73,wCS:29200,sine:1,thAx:0,allAx:2,thDist:-2}
231. r2_pboole(X4,X5,k2_pboole(X4,X6,X7)) | ~r2_pboole(X4,X5,X6) | ~v1_partfun1(k2_pboole(X4,X6,X7),X4) | ~v1_funct_1(k2_pboole(X4,X6,X7)) | ~v4_relat_1(k2_pboole(X4,X6,X7),X4) | ~v1_relat_1(k2_pboole(X4,X6,X7)) | ~v1_partfun1(X6,X4) | ~v1_funct_1(X6) | ~v4_relat_1(X6,X4) | ~v1_relat_1(X6) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~v1_partfun1(X7,X4) | ~v1_funct_1(X7) | ~v4_relat_1(X7,X4) | ~v1_relat_1(X7) [duplicate literal removal 226] {a:1,w:63,wCS:25200,sine:1,thAx:0,allAx:2,thDist:-2}
233. r2_pboole(X4,X5,k2_pboole(X4,X6,X7)) | ~r2_pboole(X4,X5,X6) | ~v1_funct_1(k2_pboole(X4,X6,X7)) | ~v4_relat_1(k2_pboole(X4,X6,X7),X4) | ~v1_relat_1(k2_pboole(X4,X6,X7)) | ~v1_partfun1(X6,X4) | ~v1_funct_1(X6) | ~v4_relat_1(X6,X4) | ~v1_relat_1(X6) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~v1_partfun1(X7,X4) | ~v1_funct_1(X7) | ~v4_relat_1(X7,X4) | ~v1_relat_1(X7) [subsumption resolution 231,81] {a:1,w:57,wCS:22800,sine:1,thAx:0,allAx:2,thDist:-2}
234. r2_pboole(X4,X5,k2_pboole(X4,X6,X7)) | ~r2_pboole(X4,X5,X6) | ~v4_relat_1(k2_pboole(X4,X6,X7),X4) | ~v1_relat_1(k2_pboole(X4,X6,X7)) | ~v1_partfun1(X6,X4) | ~v1_funct_1(X6) | ~v4_relat_1(X6,X4) | ~v1_relat_1(X6) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~v1_partfun1(X7,X4) | ~v1_funct_1(X7) | ~v4_relat_1(X7,X4) | ~v1_relat_1(X7) [subsumption resolution 233,80] {a:1,w:52,wCS:20800,sine:1,thAx:0,allAx:2,thDist:-2}
235. r2_pboole(X4,X5,k2_pboole(X4,X6,X7)) | ~r2_pboole(X4,X5,X6) | ~v1_relat_1(k2_pboole(X4,X6,X7)) | ~v1_partfun1(X6,X4) | ~v1_funct_1(X6) | ~v4_relat_1(X6,X4) | ~v1_relat_1(X6) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~v1_partfun1(X7,X4) | ~v1_funct_1(X7) | ~v4_relat_1(X7,X4) | ~v1_relat_1(X7) [subsumption resolution 234,79] {a:1,w:46,wCS:18400,sine:1,thAx:0,allAx:2,thDist:-2}
236. r2_pboole(X4,X5,k2_pboole(X4,X6,X7)) | ~r2_pboole(X4,X5,X6) | ~v1_partfun1(X6,X4) | ~v1_funct_1(X6) | ~v4_relat_1(X6,X4) | ~v1_relat_1(X6) | ~v1_partfun1(X5,X4) | ~v1_funct_1(X5) | ~v4_relat_1(X5,X4) | ~v1_relat_1(X5) | ~v1_partfun1(X7,X4) | ~v1_funct_1(X7) | ~v4_relat_1(X7,X4) | ~v1_relat_1(X7) [subsumption resolution 235,78] {a:1,w:41,wCS:16400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
257. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~1) [resolution 236,87] {a:2,w:49,wCS:4900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
266. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~1, 10) [subsumption resolution 257,132] {a:2,w:46,wCS:4600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
267. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~1, 10, 11) [subsumption resolution 266,137] {a:2,w:44,wCS:4400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
268. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) | ~v1_relat_1(sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~1, 10, 11, 12) [subsumption resolution 267,142] {a:2,w:41,wCS:4100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
269. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~v1_partfun1(sK3,sK0) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~1, 10, 11, 12, 13) [subsumption resolution 268,147] {a:2,w:39,wCS:3900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
270. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~v1_funct_1(sK3) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~1, 2, 10, 11, 12, 13) [subsumption resolution 269,92] {a:2,w:36,wCS:3600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
271. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(sK3,sK0) | ~v1_relat_1(sK3) <- (~1, 2, 3, 10, 11, 12, 13) [subsumption resolution 270,97] {a:2,w:34,wCS:3400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
272. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) | ~v1_relat_1(sK3) <- (~1, 2, 3, 4, 10, 11, 12, 13) [subsumption resolution 271,102] {a:2,w:31,wCS:3100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
273. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) | ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) | ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) | ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) <- (~1, 2, 3, 4, 5, 10, 11, 12, 13) [subsumption resolution 272,107] {a:2,w:29,wCS:2900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
275. 16 <=> v1_relat_1(k3_pboole(sK0,sK1,sK2)) [avatar definition]
277. ~v1_relat_1(k3_pboole(sK0,sK1,sK2)) <- (~16) [avatar component clause 275] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
279. 17 <=> v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) [avatar definition]
281. ~v4_relat_1(k3_pboole(sK0,sK1,sK2),sK0) <- (~17) [avatar component clause 279] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
283. 18 <=> v1_funct_1(k3_pboole(sK0,sK1,sK2)) [avatar definition]
285. ~v1_funct_1(k3_pboole(sK0,sK1,sK2)) <- (~18) [avatar component clause 283] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
287. 19 <=> v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) [avatar definition]
289. ~v1_partfun1(k3_pboole(sK0,sK1,sK2),sK0) <- (~19) [avatar component clause 287] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
291. 20 <=> r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) [avatar definition]
293. ~r2_pboole(sK0,k3_pboole(sK0,sK1,sK2),sK1) <- (~20) [avatar component clause 291] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
294. ~16 | ~17 | ~18 | ~19 | ~20 | 1 | ~2 | ~3 | ~4 | ~5 | ~10 | ~11 | ~12 | ~13 [avatar split clause 273,145,140,135,130,105,100,95,90,85,291,287,283,279,275]
326. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~16) [resolution 277,74] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
328. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, ~16) [subsumption resolution 326,112] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
329. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, ~16) [subsumption resolution 328,117] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
330. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, ~16) [subsumption resolution 329,122] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
331. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, ~16) [subsumption resolution 330,127] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
332. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, ~16) [subsumption resolution 331,132] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
333. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, ~16) [subsumption resolution 332,137] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
334. ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, 12, ~16) [subsumption resolution 333,142] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
335. $false <- (6, 7, 8, 9, 10, 11, 12, 13, ~16) [subsumption resolution 334,147] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
336. ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | 16 [avatar contradiction clause 335]
338. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~17) [resolution 281,75] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
340. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, ~17) [subsumption resolution 338,112] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
341. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, ~17) [subsumption resolution 340,117] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
342. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, ~17) [subsumption resolution 341,122] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
343. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, ~17) [subsumption resolution 342,127] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
344. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, ~17) [subsumption resolution 343,132] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
345. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, ~17) [subsumption resolution 344,137] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
346. ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, 12, ~17) [subsumption resolution 345,142] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
347. $false <- (6, 7, 8, 9, 10, 11, 12, 13, ~17) [subsumption resolution 346,147] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
348. ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | 17 [avatar contradiction clause 347]
350. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~18) [resolution 285,76] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
352. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, ~18) [subsumption resolution 350,112] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
353. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, ~18) [subsumption resolution 352,117] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
354. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, ~18) [subsumption resolution 353,122] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
355. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, ~18) [subsumption resolution 354,127] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
356. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, ~18) [subsumption resolution 355,132] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
357. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, ~18) [subsumption resolution 356,137] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
358. ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, 12, ~18) [subsumption resolution 357,142] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
359. $false <- (6, 7, 8, 9, 10, 11, 12, 13, ~18) [subsumption resolution 358,147] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
360. ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | 18 [avatar contradiction clause 359]
362. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~19) [resolution 289,77] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
364. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, ~19) [subsumption resolution 362,112] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
365. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, ~19) [subsumption resolution 364,117] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
366. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, ~19) [subsumption resolution 365,122] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
367. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, ~19) [subsumption resolution 366,127] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
368. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, ~19) [subsumption resolution 367,132] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
369. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, ~19) [subsumption resolution 368,137] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
370. ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, 12, ~19) [subsumption resolution 369,142] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
371. $false <- (6, 7, 8, 9, 10, 11, 12, 13, ~19) [subsumption resolution 370,147] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
372. ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | 19 [avatar contradiction clause 371]
374. ~v1_partfun1(sK2,sK0) | ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~20) [resolution 293,63] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
376. ~v1_funct_1(sK2) | ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, ~20) [subsumption resolution 374,112] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
377. ~v4_relat_1(sK2,sK0) | ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, ~20) [subsumption resolution 376,117] {a:3,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
378. ~v1_relat_1(sK2) | ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, ~20) [subsumption resolution 377,122] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
379. ~v1_partfun1(sK1,sK0) | ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, ~20) [subsumption resolution 378,127] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
380. ~v1_funct_1(sK1) | ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, ~20) [subsumption resolution 379,132] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
381. ~v4_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, ~20) [subsumption resolution 380,137] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
382. ~v1_relat_1(sK1) <- (6, 7, 8, 9, 10, 11, 12, ~20) [subsumption resolution 381,142] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
383. $false <- (6, 7, 8, 9, 10, 11, 12, 13, ~20) [subsumption resolution 382,147] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
384. ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | 20 [avatar contradiction clause 383]
385. $false [avatar sat refutation 88,93,98,103,108,113,118,123,128,133,138,143,148,294,336,348,360,372,384] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:321,thDist:-321}
% SZS output end Proof for pboole__t24_pboole
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 13
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 35
% Generated clauses: 254
% Active clauses: 48
% Passive clauses: 59
% Final active clauses: 44
% Final passive clauses: 9
% Main loop iterations started: 50
% 
% >>> Simplifying Inferences
% Duplicate literals: 368
% Fw subsumption resolutions: 84
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Forward subsumptions: 27
% 
% >>> Generating Inferences
% Binary resolution: 15
% Unit resulting resolution: 5
% Forward superposition: 18
% Backward superposition: 24
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 5
% Sat splitting refutations: 10
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 28
% SAT solver unit clauses: 15
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.003 s

