i: [1,54,0,2,1,1,0,0] t38_wellord1
a: [3,87,0,2,1,54]
s: 87
i: [1,55,0,3,1,1,0,0] t38_wellord1
a: [3,82,0,3,1,55]
s: 82
i: [1,72,0,2,1,0,0,255] rc3_relat_1
a: [3,92,0,2,1,72]
s: 92
i: [1,74,0,2,1,0,0,255] rc6_funct_1
a: [3,97,0,2,1,74]
s: 97
i: [1,73,0,2,1,0,0,255] rc6_funct_1
a: [3,102,0,2,1,73]
s: 102
i: [1,76,0,2,1,0,0,2] rc1_funct_1
a: [3,107,0,2,1,76]
s: 107
i: [1,75,0,2,1,0,0,2] rc1_funct_1
a: [3,112,0,2,1,75]
s: 112
i: [1,78,0,2,1,0,0,255] rc2_funct_1
a: [3,117,0,2,1,78]
s: 117
i: [1,77,0,2,1,0,0,255] rc2_funct_1
a: [3,122,0,2,1,77]
s: 122
i: [1,56,0,3,1,0,0,3] dt_k4_relat_1
s: 56
i: [1,59,0,3,1,0,0,255] fc4_funct_1
s: 59
i: [1,60,0,3,1,0,0,2] fc3_funct_1
s: 60
i: [1,61,0,3,1,0,0,2] fc3_funct_1
s: 61
i: [1,62,0,3,1,0,0,255] fc28_relat_1
s: 62
i: [1,57,0,5,1,0,0,4] rd2_relat_1
s: 57
i: [1,58,0,5,1,0,0,4] rd1_relat_1
s: 58
i: [1,69,0,5,1,0,0,4] idempotence_k2_xboole_0
s: 69
i: [1,70,0,7,1,0,0,4] commutativity_k2_xboole_0
s: 70
i: [1,63,0,8,2,0,0,2] t37_wellord1
s: 63
i: [1,71,0,8,3,0,0,4] fc3_relat_1
s: 71
i: [1,64,0,10,2,0,0,3] d6_relat_1
s: 64
d: [2,128,1,12,2,0,44,64,58]
d: [2,131,1,9,1,0,21,128,56]
d: [2,132,1,7,1,0,22,131,57]
d: [2,133,1,5,1,0,22,132,69]
s: 133
i: [1,65,0,11,4,0,0,1] d8_wellord1
s: 65
d: [2,129,1,12,2,0,44,64,57]
d: [2,134,1,9,1,0,21,129,56]
d: [2,135,1,7,1,0,22,134,58]
d: [2,136,1,5,1,0,22,135,69]
s: 136
i: [1,66,0,11,4,0,0,1] d8_wellord1
s: 66
i: [1,67,0,13,4,0,0,1] d8_wellord1
s: 67
d: [2,137,2,10,2,0,44,63,133]
d: [2,138,2,7,1,0,21,137,56]
s: 138
i: [1,68,0,15,6,0,0,1] d8_wellord1
s: 68
d: [2,139,1,17,6,0,40,68,63]
d: [2,144,1,13,4,0,18,139]
d: [2,145,1,9,3,0,21,144,61]
d: [2,146,1,5,2,0,21,145,56]
s: 146
d: [2,149,2,0,0,2,52,87,82,146]
e: 149
d: [2,150,2,2,1,1,40,146,82]
d: [2,152,2,0,0,2,21,150,87]
e: 152
f: 152,54,55
% Refutation found. Thanks to Tanya!
% SZS status Theorem for wellord1__t38_wellord1
% SZS output start Proof for wellord1__t38_wellord1
1. ! [X0] : (v1_relat_1(X0) => r4_wellord1(X0,X0)) [input t38_wellord1]
2. ~! [X0] : (v1_relat_1(X0) => r4_wellord1(X0,X0)) [negated conjecture 1]
5. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r4_wellord1(X0,X1) <=> ? [X2] : (r3_wellord1(X0,X1,X2) & v1_funct_1(X2) & v1_relat_1(X2))))) [input d8_wellord1]
9. ! [X0] : v1_relat_1(k4_relat_1(X0)) [input dt_k4_relat_1]
12. ! [X0] : (v1_funct_1(k4_relat_1(X0)) & v1_relat_1(k4_relat_1(X0))) [input fc3_funct_1]
22. ! [X0] : (v1_relat_1(X0) => r3_wellord1(X0,X0,k4_relat_1(k1_relat_1(X0)))) [input t37_wellord1]
32. ? [X0] : (~r4_wellord1(X0,X0) & v1_relat_1(X0)) [ennf transformation 2]
33. ! [X0] : (r3_wellord1(X0,X0,k4_relat_1(k1_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 22]
35. ! [X0] : (! [X1] : ((r4_wellord1(X0,X1) <=> ? [X2] : (r3_wellord1(X0,X1,X2) & v1_funct_1(X2) & v1_relat_1(X2))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 5]
38. ? [X0] : (~r4_wellord1(X0,X0) & v1_relat_1(X0)) => (~r4_wellord1(sK0,sK0) & v1_relat_1(sK0)) [choice axiom]
39. ~r4_wellord1(sK0,sK0) & v1_relat_1(sK0) [skolemisation 32,38]
40. ! [X0] : (! [X1] : (((r4_wellord1(X0,X1) | ! [X2] : (~r3_wellord1(X0,X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X2] : (r3_wellord1(X0,X1,X2) & v1_funct_1(X2) & v1_relat_1(X2)) | ~r4_wellord1(X0,X1))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [nnf transformation 35]
41. ! [X0] : (! [X1] : (((r4_wellord1(X0,X1) | ! [X2] : (~r3_wellord1(X0,X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & (? [X3] : (r3_wellord1(X0,X1,X3) & v1_funct_1(X3) & v1_relat_1(X3)) | ~r4_wellord1(X0,X1))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [rectify 40]
42. ! [X1,X0] : (? [X3] : (r3_wellord1(X0,X1,X3) & v1_funct_1(X3) & v1_relat_1(X3)) => (r3_wellord1(X0,X1,sK1(X0,X1)) & v1_funct_1(sK1(X0,X1)) & v1_relat_1(sK1(X0,X1)))) [choice axiom]
43. ! [X0] : (! [X1] : (((r4_wellord1(X0,X1) | ! [X2] : (~r3_wellord1(X0,X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) & ((r3_wellord1(X0,X1,sK1(X0,X1)) & v1_funct_1(sK1(X0,X1)) & v1_relat_1(sK1(X0,X1))) | ~r4_wellord1(X0,X1))) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [skolemisation 41,42]
54. v1_relat_1(sK0) [cnf transformation 39] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. ~r4_wellord1(sK0,sK0) [cnf transformation 39] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
56. v1_relat_1(k4_relat_1(X0)) [cnf transformation 9] {a:0,w:3,wCS:1200,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
61. v1_funct_1(k4_relat_1(X0)) [cnf transformation 12] {a:0,w:3,wCS:1200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
63. r3_wellord1(X0,X0,k4_relat_1(k1_relat_1(X0))) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
68. ~r3_wellord1(X0,X1,X2) | r4_wellord1(X0,X1) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 43] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
80. 1 <=> r4_wellord1(sK0,sK0) [avatar definition]
82. ~r4_wellord1(sK0,sK0) <- (~1) [avatar component clause 80] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. ~1 [avatar split clause 55,80]
85. 2 <=> v1_relat_1(sK0) [avatar definition]
87. v1_relat_1(sK0) <- (2) [avatar component clause 85] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
88. 2 [avatar split clause 54,85]
139. r4_wellord1(X0,X0) | ~v1_funct_1(k4_relat_1(k1_relat_1(X0))) | ~v1_relat_1(k4_relat_1(k1_relat_1(X0))) | ~v1_relat_1(X0) | ~v1_relat_1(X0) | ~v1_relat_1(X0) [resolution 68,63] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:3,thDist:-3}
144. r4_wellord1(X0,X0) | ~v1_funct_1(k4_relat_1(k1_relat_1(X0))) | ~v1_relat_1(k4_relat_1(k1_relat_1(X0))) | ~v1_relat_1(X0) [duplicate literal removal 139] {a:1,w:13,wCS:5200,sine:1,thAx:0,allAx:3,thDist:-3}
145. r4_wellord1(X0,X0) | ~v1_relat_1(k4_relat_1(k1_relat_1(X0))) | ~v1_relat_1(X0) [subsumption resolution 144,61] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:3,thDist:-3}
146. r4_wellord1(X0,X0) | ~v1_relat_1(X0) [subsumption resolution 145,56] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
150. ~v1_relat_1(sK0) <- (~1) [resolution 146,82] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
152. $false <- (~1, 2) [subsumption resolution 150,87] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
153. 1 | ~2 [avatar contradiction clause 152]
154. $false [avatar sat refutation 83,88,153] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:39,thDist:-39}
% SZS output end Proof for wellord1__t38_wellord1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 3
% 
% >>> Saturation
% Initial clauses: 25
% Generated clauses: 62
% Active clauses: 25
% Passive clauses: 31
% Final active clauses: 25
% Final passive clauses: 2
% Main loop iterations started: 29
% 
% >>> Simplifying Inferences
% Duplicate literals: 7
% Fw subsumption resolutions: 8
% Fw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% Forward superposition: 5
% Backward superposition: 2
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 9
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

