i: [1,26,0,3,1,1,0,0] l23_arytm_1
a: [3,41,0,3,1,26]
s: 41
i: [1,25,0,3,1,1,0,0] l23_arytm_1
a: [3,46,0,3,1,25]
s: 46
i: [1,24,0,3,1,1,0,0] l23_arytm_1
a: [3,51,0,3,1,24]
s: 51
i: [1,27,0,11,1,1,0,0] l23_arytm_1
a: [3,36,0,11,1,27]
s: 36
i: [1,29,0,4,1,0,0,1] existence_m1_subset_1
s: 29
i: [1,30,0,11,3,0,0,2] dt_k7_arytm_2
s: 30
i: [1,31,0,11,3,0,0,1] dt_k1_arytm_1
s: 31
i: [1,32,0,13,3,0,0,1] commutativity_k7_arytm_2
s: 32
i: [1,28,0,20,4,0,0,1] l22_arytm_1
s: 28
d: [2,58,1,20,4,1,44,36,28]
d: [2,63,1,17,3,2,21,58,46]
d: [2,64,1,14,2,3,21,63,41]
d: [2,65,1,11,1,4,21,64,51]
a: [3,69,1,11,1,65]
s: 69
d: [2,59,1,24,6,0,44,31,28]
d: [2,60,1,21,5,0,18,59]
d: [2,71,1,16,4,0,21,60,31]
s: 71
d: [2,72,2,20,4,1,44,69,28]
d: [2,73,2,17,3,2,21,72,41]
d: [2,74,2,14,2,3,21,73,46]
d: [2,75,2,11,1,4,21,74,51]
a: [3,79,2,11,1,75]
s: 79
d: [2,86,3,17,3,1,44,79,32]
d: [2,89,3,6,2,1,19,86]
d: [2,90,3,3,1,2,21,89,41]
d: [2,91,3,0,0,3,21,90,46]
e: 91
d: [2,87,3,17,3,1,44,79,32]
d: [2,88,3,6,2,1,19,87]
d: [2,93,3,3,1,2,21,88,46]
d: [2,94,3,0,0,3,21,93,41]
e: 94
f: 94,75,65,24,25,26,27
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__l23_arytm_1
% SZS output start Proof for arytm_1__l23_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => k1_arytm_1(k1_arytm_1(X0,X1),X2) = k1_arytm_1(k1_arytm_1(X0,X2),X1)))) [input l23_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => k1_arytm_1(k1_arytm_1(X0,X1),X2) = k1_arytm_1(k1_arytm_1(X0,X2),X1)))) [negated conjecture 1]
3. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => k7_arytm_2(X0,X1) = k7_arytm_2(X1,X0)) [input commutativity_k7_arytm_2]
9. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => k1_arytm_1(k1_arytm_1(X0,X1),X2) = k1_arytm_1(X0,k7_arytm_2(X2,X1))))) [input l22_arytm_1]
10. ? [X0] : (? [X1] : (? [X2] : (k1_arytm_1(k1_arytm_1(X0,X1),X2) != k1_arytm_1(k1_arytm_1(X0,X2),X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
11. ! [X0] : (! [X1] : (! [X2] : (k1_arytm_1(k1_arytm_1(X0,X1),X2) = k1_arytm_1(X0,k7_arytm_2(X2,X1)) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 9]
16. ! [X0,X1] : (k7_arytm_2(X0,X1) = k7_arytm_2(X1,X0) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 3]
17. ! [X0,X1] : (k7_arytm_2(X0,X1) = k7_arytm_2(X1,X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 16]
18. ? [X0] : (? [X1] : (? [X2] : (k1_arytm_1(k1_arytm_1(X0,X1),X2) != k1_arytm_1(k1_arytm_1(X0,X2),X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (? [X2] : (k1_arytm_1(k1_arytm_1(sK0,X1),X2) != k1_arytm_1(k1_arytm_1(sK0,X2),X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
19. ? [X1] : (? [X2] : (k1_arytm_1(k1_arytm_1(sK0,X1),X2) != k1_arytm_1(k1_arytm_1(sK0,X2),X1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(X1,k2_arytm_2)) => (? [X2] : (k1_arytm_1(k1_arytm_1(sK0,sK1),X2) != k1_arytm_1(k1_arytm_1(sK0,X2),sK1) & m1_subset_1(X2,k2_arytm_2)) & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
20. ? [X2] : (k1_arytm_1(k1_arytm_1(sK0,sK1),X2) != k1_arytm_1(k1_arytm_1(sK0,X2),sK1) & m1_subset_1(X2,k2_arytm_2)) => (k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) != k1_arytm_1(k1_arytm_1(sK0,sK2),sK1) & m1_subset_1(sK2,k2_arytm_2)) [choice axiom]
21. ((k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) != k1_arytm_1(k1_arytm_1(sK0,sK2),sK1) & m1_subset_1(sK2,k2_arytm_2)) & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 10,20,19,18]
24. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 21] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
25. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 21] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
26. m1_subset_1(sK2,k2_arytm_2) [cnf transformation 21] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
27. k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) != k1_arytm_1(k1_arytm_1(sK0,sK2),sK1) [cnf transformation 21] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
28. k1_arytm_1(k1_arytm_1(X0,X1),X2) = k1_arytm_1(X0,k7_arytm_2(X2,X1)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 11] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
32. k7_arytm_2(X0,X1) = k7_arytm_2(X1,X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 17] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. 1 <=> k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) = k1_arytm_1(k1_arytm_1(sK0,sK2),sK1) [avatar definition]
36. k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) != k1_arytm_1(k1_arytm_1(sK0,sK2),sK1) <- (~1) [avatar component clause 34] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
37. ~1 [avatar split clause 27,34]
39. 2 <=> m1_subset_1(sK2,k2_arytm_2) [avatar definition]
41. m1_subset_1(sK2,k2_arytm_2) <- (2) [avatar component clause 39] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
42. 2 [avatar split clause 26,39]
44. 3 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
46. m1_subset_1(sK1,k2_arytm_2) <- (3) [avatar component clause 44] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
47. 3 [avatar split clause 25,44]
49. 4 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
51. m1_subset_1(sK0,k2_arytm_2) <- (4) [avatar component clause 49] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
52. 4 [avatar split clause 24,49]
58. k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) != k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1) [superposition 36,28] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
63. k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) != k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1, 3) [subsumption resolution 58,46] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
64. k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) != k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) | ~m1_subset_1(sK0,k2_arytm_2) <- (~1, 2, 3) [subsumption resolution 63,41] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
65. k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) != k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) <- (~1, 2, 3, 4) [subsumption resolution 64,51] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
67. 5 <=> k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) = k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) [avatar definition]
69. k1_arytm_1(k1_arytm_1(sK0,sK1),sK2) != k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) <- (~5) [avatar component clause 67] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
70. ~5 | 1 | ~2 | ~3 | ~4 [avatar split clause 65,49,44,39,34,67]
72. k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) != k1_arytm_1(sK0,k7_arytm_2(sK2,sK1)) | ~m1_subset_1(sK2,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (~5) [superposition 69,28] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
73. k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) != k1_arytm_1(sK0,k7_arytm_2(sK2,sK1)) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (2, ~5) [subsumption resolution 72,41] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
74. k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) != k1_arytm_1(sK0,k7_arytm_2(sK2,sK1)) | ~m1_subset_1(sK0,k2_arytm_2) <- (2, 3, ~5) [subsumption resolution 73,46] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
75. k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) != k1_arytm_1(sK0,k7_arytm_2(sK2,sK1)) <- (2, 3, 4, ~5) [subsumption resolution 74,51] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
77. 6 <=> k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) = k1_arytm_1(sK0,k7_arytm_2(sK2,sK1)) [avatar definition]
79. k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) != k1_arytm_1(sK0,k7_arytm_2(sK2,sK1)) <- (~6) [avatar component clause 77] {a:2,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
80. ~6 | ~2 | ~3 | ~4 | 5 [avatar split clause 75,67,49,44,39,77]
87. k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) != k1_arytm_1(sK0,k7_arytm_2(sK1,sK2)) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK2,k2_arytm_2) <- (~6) [superposition 79,32] {a:3,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
88. ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK2,k2_arytm_2) <- (~6) [trivial inequality removal 87] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
93. ~m1_subset_1(sK2,k2_arytm_2) <- (3, ~6) [subsumption resolution 88,46] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
94. $false <- (2, 3, ~6) [subsumption resolution 93,41] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
95. ~2 | ~3 | 6 [avatar contradiction clause 94]
96. $false [avatar sat refutation 37,42,47,52,70,80,95] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:55,thDist:-55}
% SZS output end Proof for arytm_1__l23_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 8
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 9
% Generated clauses: 46
% Active clauses: 12
% Passive clauses: 13
% Final active clauses: 12
% Final passive clauses: 1
% Main loop iterations started: 12
% 
% >>> Simplifying Inferences
% Duplicate literals: 10
% Trivial inequalities: 2
% Fw subsumption resolutions: 12
% 
% >>> Deletion Inferences
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Forward superposition: 6
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

