i: [1,76,0,2,1,1,0,0] t87_rfunct_1
a: [3,142,0,2,1,76]
s: 142
i: [1,75,0,2,1,1,0,0] t87_rfunct_1
a: [3,147,0,2,1,75]
s: 147
i: [1,74,0,2,1,1,0,0] t87_rfunct_1
a: [3,152,0,2,1,74]
s: 152
i: [1,78,0,4,1,1,0,0] t87_rfunct_1
a: [3,132,0,4,1,78]
s: 132
i: [1,77,0,4,1,1,0,0] t87_rfunct_1
a: [3,137,0,4,1,77]
s: 137
i: [1,79,0,6,1,1,0,0] t87_rfunct_1
a: [3,127,0,6,1,79]
s: 127
i: [1,123,0,2,1,0,0,1] rc1_funct_1
a: [3,157,0,2,1,123]
s: 157
i: [1,122,0,2,1,0,0,1] rc1_funct_1
a: [3,162,0,2,1,122]
s: 162
i: [1,98,0,5,1,0,0,255] idempotence_k2_xboole_0
s: 98
i: [1,93,0,6,3,0,0,255] cc5_valued_0
s: 93
i: [1,95,0,6,3,0,0,255] cc6_valued_0
s: 95
d: [2,164,1,4,2,1,40,93,142]
d: [2,165,1,2,1,2,21,164,152]
a: [3,169,1,2,1,165]
s: 169
i: [1,100,0,6,2,0,0,1] dt_k5_relat_1
s: 100
d: [2,171,1,4,2,1,40,95,142]
d: [2,172,1,2,1,2,21,171,152]
a: [3,176,1,2,1,172]
s: 176
i: [1,99,0,7,1,0,0,255] commutativity_k2_xboole_0
s: 99
i: [1,104,0,8,3,0,0,255] fc43_valued_0
s: 104
i: [1,105,0,8,3,0,0,255] fc43_valued_0
s: 105
i: [1,106,0,8,3,0,0,255] fc44_valued_0
s: 106
i: [1,107,0,8,3,0,0,255] fc44_valued_0
s: 107
i: [1,116,0,8,3,0,0,1] fc8_funct_1
s: 116
i: [1,117,0,8,3,0,0,1] fc8_funct_1
s: 117
i: [1,119,0,8,3,0,0,1] fc45_valued_0
s: 119
i: [1,120,0,8,3,0,0,1] fc45_valued_0
s: 120
i: [1,121,0,8,3,0,0,1] fc3_relat_1
s: 121
i: [1,83,0,10,5,0,0,1] cc1_seq_2
s: 83
i: [1,84,0,10,5,0,0,1] cc1_seq_2
s: 84
d: [2,186,1,16,4,1,40,83,132]
a: [3,191,1,4,1,186]
s: 191
d: [2,234,2,0,0,2,52,152,191,100]
e: 234
d: [2,235,2,2,1,1,40,191,100]
d: [2,237,2,0,0,2,21,235,152]
e: 237
r: 191
i: [1,89,0,10,5,0,0,1] d8_seq_2
s: 89
a: [3,190,1,4,1,186]
s: 190
i: [1,90,0,10,5,0,0,1] d8_seq_2
s: 90
i: [1,101,0,11,2,0,0,255] rd5_relat_1
s: 101
a: [3,195,1,4,1,186]
s: 195
d: [2,245,2,0,0,3,52,152,147,195,117]
e: 245
d: [2,246,2,4,2,1,40,195,117]
d: [2,248,2,2,1,2,21,246,147]
d: [2,249,2,0,0,3,21,248,152]
e: 249
r: 195
i: [1,88,0,12,6,0,0,1] cc2_seq_2
s: 88
a: [3,194,1,4,1,186]
s: 194
i: [1,91,0,12,6,0,0,1] d8_seq_2
s: 91
i: [1,102,0,12,5,0,0,255] fc8_valued_0
s: 102
a: [3,199,1,4,1,186]
s: 199
d: [2,255,2,0,0,3,52,152,142,199,120]
e: 255
d: [2,256,2,4,2,1,40,199,120]
d: [2,258,2,2,1,2,21,256,142]
d: [2,259,2,0,0,3,21,258,152]
e: 259
r: 199
i: [1,103,0,12,5,0,0,255] fc7_valued_0
s: 103
a: [3,198,1,4,1,186]
s: 198
i: [1,108,0,12,5,0,0,2] fc7_rfunct_1
s: 108
i: [1,109,0,12,5,0,0,2] fc7_rfunct_1
s: 109
a: [3,203,1,4,1,186]
s: 203
i: [1,110,0,12,5,0,0,2] fc7_rfunct_1
s: 110
d: [2,187,1,16,4,1,40,83,137]
a: [3,216,1,4,1,187]
s: 216
d: [2,279,2,0,0,3,52,152,142,216,120]
e: 279
d: [2,280,2,4,2,1,40,216,120]
d: [2,282,2,2,1,2,21,280,142]
d: [2,283,2,0,0,3,21,282,152]
e: 283
r: 216
i: [1,111,0,12,5,0,0,2] fc7_rfunct_1
s: 111
i: [1,112,0,12,5,0,0,2] fc6_rfunct_1
s: 112
a: [3,208,1,4,1,187]
s: 208
d: [2,290,2,0,0,2,52,152,208,100]
e: 290
d: [2,291,2,2,1,1,40,208,100]
d: [2,293,2,0,0,2,21,291,152]
e: 293
r: 208
i: [1,113,0,12,5,0,0,2] fc6_rfunct_1
s: 113
a: [3,207,1,4,1,187]
s: 207
i: [1,114,0,12,5,0,0,2] fc6_rfunct_1
s: 114
i: [1,115,0,12,5,0,0,2] fc6_rfunct_1
s: 115
a: [3,212,1,4,1,187]
s: 212
d: [2,296,2,0,0,3,52,152,147,212,117]
e: 296
d: [2,297,2,4,2,1,40,212,117]
d: [2,299,2,2,1,2,21,297,147]
d: [2,300,2,0,0,3,21,299,152]
e: 300
r: 212
i: [1,118,0,12,5,0,0,1] fc9_valued_0
s: 118
a: [3,211,1,4,1,187]
s: 211
i: [1,96,0,20,6,0,0,2] t85_rfunct_1
s: 96
i: [1,97,0,20,6,0,0,2] t86_rfunct_1
s: 97
a: [3,215,1,4,1,187]
s: 215
a: [3,220,1,4,1,187]
s: 220
d: [2,222,1,16,4,1,40,84,132]
a: [3,227,1,4,1,222]
s: 227
d: [2,223,1,16,4,1,40,84,137]
a: [3,232,1,4,1,223]
s: 232
d: [2,178,1,12,4,0,40,120,95]
d: [2,180,1,8,3,0,21,178,100]
s: 180
d: [2,265,2,8,2,1,40,198,95]
d: [2,267,2,4,1,2,21,265,190]
a: [3,271,2,4,1,267]
s: 271
d: [2,179,1,12,4,0,40,120,93]
d: [2,181,1,8,3,0,21,179,100]
s: 181
d: [2,266,2,8,2,1,40,198,93]
d: [2,273,2,4,1,2,21,266,190]
a: [3,277,2,4,1,273]
s: 277
d: [2,302,1,16,6,0,40,118,95]
d: [2,308,1,12,5,0,21,302,121]
s: 308
d: [2,303,1,16,6,0,40,118,93]
d: [2,309,1,12,5,0,21,303,121]
s: 309
d: [2,326,2,8,2,1,40,215,95]
d: [2,328,2,4,1,2,21,326,207]
a: [3,332,2,4,1,328]
s: 332
d: [2,285,1,28,9,0,40,111,88]
d: [2,287,1,24,8,0,21,285,120]
d: [2,288,1,20,7,0,21,287,117]
d: [2,289,1,16,6,0,21,288,100]
s: 289
d: [2,327,2,8,2,1,40,215,93]
d: [2,334,2,4,1,2,21,327,207]
a: [3,338,2,4,1,334]
s: 338
d: [2,316,1,44,10,0,40,97,88]
d: [2,322,1,38,9,0,21,316,120]
d: [2,323,1,32,8,0,21,322,117]
d: [2,324,1,26,7,0,21,323,100]
s: 324
d: [2,313,1,42,7,0,44,96,101]
d: [2,315,1,36,6,0,21,313,100]
s: 315
d: [2,354,2,20,9,0,40,289,115]
d: [2,358,2,14,6,0,18,354]
s: 358
d: [2,320,1,42,7,0,44,97,101]
d: [2,325,1,36,6,0,21,320,100]
s: 325
d: [2,379,3,10,5,1,40,358,127]
d: [2,383,3,8,4,2,21,379,147]
d: [2,384,3,6,3,3,21,383,152]
d: [2,385,3,4,2,4,21,384,142]
a: [3,389,3,2,1,385]
s: 389
d: [2,361,2,22,7,0,44,324,98]
d: [2,365,2,18,6,0,18,361]
s: 365
d: [2,371,2,26,6,0,44,315,98]
d: [2,375,2,20,5,0,18,371]
s: 375
d: [2,397,2,26,6,0,44,325,98]
d: [2,401,2,20,5,0,18,397]
s: 401
d: [2,355,2,28,10,0,40,289,96]
d: [2,357,2,22,7,0,18,355]
s: 357
a: [3,393,3,2,1,385]
s: 393
d: [2,360,2,28,10,0,40,324,115]
d: [2,366,2,22,7,0,18,360]
s: 366
d: [2,359,2,34,11,0,40,324,96]
d: [2,367,2,28,8,0,18,359]
s: 367
d: [2,450,3,0,0,8,52,142,152,147,220,203,227,232,127,367]
e: 450
d: [2,451,3,22,7,1,40,367,127]
d: [2,460,3,20,6,2,21,451,142]
d: [2,461,3,18,5,3,21,460,147]
d: [2,462,3,16,4,4,21,461,152]
d: [2,463,3,12,3,5,21,462,232]
d: [2,464,3,8,2,6,21,463,227]
d: [2,465,3,4,1,7,21,464,203]
d: [2,466,3,0,0,8,21,465,220]
e: 466
f: 466,300,293,283,259,249,237,223,222,187,186,74,75,76,77,78,79
% Refutation found. Thanks to Tanya!
% SZS status Theorem for rfunct_1__t87_rfunct_1
% SZS output start Proof for rfunct_1__t87_rfunct_1
1. ! [X0] : ! [X1] : ! [X2] : ((v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((v1_comseq_2(k5_relat_1(X2,X1)) & v1_comseq_2(k5_relat_1(X2,X0))) => v1_comseq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))))) [input t87_rfunct_1]
2. ~! [X0] : ! [X1] : ! [X2] : ((v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((v1_comseq_2(k5_relat_1(X2,X1)) & v1_comseq_2(k5_relat_1(X2,X0))) => v1_comseq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))))) [negated conjecture 1]
3. ! [X0] : ((v1_comseq_2(X0) & v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v2_seq_2(X0) & v1_seq_2(X0) & v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0))) [input cc1_seq_2]
4. ! [X0] : ((v2_seq_2(X0) & v1_seq_2(X0) & v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v1_comseq_2(X0) & v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0))) [input cc2_seq_2]
10. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k5_relat_1(X0,X1))) [input dt_k5_relat_1]
14. ! [X0,X1] : ((v3_valued_0(X0) & v1_relat_1(X0)) => (v3_valued_0(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1)))) [input fc45_valued_0]
18. ! [X0,X1] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1)))) [input fc8_funct_1]
24. ! [X0] : ! [X1] : ! [X2] : ((v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((v1_seq_2(k5_relat_1(X2,X1)) & v1_seq_2(k5_relat_1(X2,X0))) => v1_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))))) [input t85_rfunct_1]
25. ! [X0] : ! [X1] : ! [X2] : ((v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((v2_seq_2(k5_relat_1(X2,X1)) & v2_seq_2(k5_relat_1(X2,X0))) => v2_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))))) [input t86_rfunct_1]
26. ~! [X0,X1,X2] : ((v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((v1_comseq_2(k5_relat_1(X2,X1)) & v1_comseq_2(k5_relat_1(X2,X0))) => v1_comseq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))))) [flattening 2]
27. ! [X0,X1,X2] : ((v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((v1_seq_2(k5_relat_1(X2,X1)) & v1_seq_2(k5_relat_1(X2,X0))) => v1_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))))) [flattening 24]
28. ! [X0,X1,X2] : ((v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((v2_seq_2(k5_relat_1(X2,X1)) & v2_seq_2(k5_relat_1(X2,X0))) => v2_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))))) [flattening 25]
30. ? [X0,X1,X2] : ((~v1_comseq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))) & (v1_comseq_2(k5_relat_1(X2,X1)) & v1_comseq_2(k5_relat_1(X2,X0)))) & (v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2))) [ennf transformation 26]
31. ? [X0,X1,X2] : (~v1_comseq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))) & v1_comseq_2(k5_relat_1(X2,X1)) & v1_comseq_2(k5_relat_1(X2,X0)) & v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) [flattening 30]
32. ! [X0] : ((v2_seq_2(X0) & v1_seq_2(X0) & v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) | (~v1_comseq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 3]
33. ! [X0] : ((v2_seq_2(X0) & v1_seq_2(X0) & v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) | ~v1_comseq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 32]
34. ! [X0] : ((v1_comseq_2(X0) & v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) | (~v2_seq_2(X0) | ~v1_seq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 4]
35. ! [X0] : ((v1_comseq_2(X0) & v3_valued_0(X0) & v1_funct_1(X0) & v1_relat_1(X0)) | ~v2_seq_2(X0) | ~v1_seq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 34]
42. ! [X0,X1,X2] : ((v1_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))) | (~v1_seq_2(k5_relat_1(X2,X1)) | ~v1_seq_2(k5_relat_1(X2,X0)))) | (~v3_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) [ennf transformation 27]
43. ! [X0,X1,X2] : (v1_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))) | ~v1_seq_2(k5_relat_1(X2,X1)) | ~v1_seq_2(k5_relat_1(X2,X0)) | ~v3_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) [flattening 42]
44. ! [X0,X1,X2] : ((v2_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))) | (~v2_seq_2(k5_relat_1(X2,X1)) | ~v2_seq_2(k5_relat_1(X2,X0)))) | (~v3_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) [ennf transformation 28]
45. ! [X0,X1,X2] : (v2_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))) | ~v2_seq_2(k5_relat_1(X2,X1)) | ~v2_seq_2(k5_relat_1(X2,X0)) | ~v3_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) [flattening 44]
46. ! [X0,X1] : (v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 10]
60. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 18]
61. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 60]
64. ! [X0,X1] : ((v3_valued_0(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | (~v3_valued_0(X0) | ~v1_relat_1(X0))) [ennf transformation 14]
65. ! [X0,X1] : ((v3_valued_0(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | ~v3_valued_0(X0) | ~v1_relat_1(X0)) [flattening 64]
68. ? [X0,X1,X2] : (~v1_comseq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))) & v1_comseq_2(k5_relat_1(X2,X1)) & v1_comseq_2(k5_relat_1(X2,X0)) & v3_valued_0(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => (~v1_comseq_2(k5_relat_1(sK2,k2_xboole_0(sK0,sK1))) & v1_comseq_2(k5_relat_1(sK2,sK1)) & v1_comseq_2(k5_relat_1(sK2,sK0)) & v3_valued_0(sK2) & v1_funct_1(sK2) & v1_relat_1(sK2)) [choice axiom]
69. ~v1_comseq_2(k5_relat_1(sK2,k2_xboole_0(sK0,sK1))) & v1_comseq_2(k5_relat_1(sK2,sK1)) & v1_comseq_2(k5_relat_1(sK2,sK0)) & v3_valued_0(sK2) & v1_funct_1(sK2) & v1_relat_1(sK2) [skolemisation 31,68]
74. v1_relat_1(sK2) [cnf transformation 69] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
75. v1_funct_1(sK2) [cnf transformation 69] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
76. v3_valued_0(sK2) [cnf transformation 69] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
77. v1_comseq_2(k5_relat_1(sK2,sK0)) [cnf transformation 69] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
78. v1_comseq_2(k5_relat_1(sK2,sK1)) [cnf transformation 69] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
79. ~v1_comseq_2(k5_relat_1(sK2,k2_xboole_0(sK0,sK1))) [cnf transformation 69] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
83. ~v1_comseq_2(X0) | v1_seq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
84. ~v1_comseq_2(X0) | v2_seq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
88. ~v2_seq_2(X0) | v1_comseq_2(X0) | ~v1_seq_2(X0) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
96. v1_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))) | ~v1_seq_2(k5_relat_1(X2,X1)) | ~v1_seq_2(k5_relat_1(X2,X0)) | ~v3_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [cnf transformation 43] {a:0,w:20,wCS:8000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
97. v2_seq_2(k5_relat_1(X2,k2_xboole_0(X0,X1))) | ~v2_seq_2(k5_relat_1(X2,X1)) | ~v2_seq_2(k5_relat_1(X2,X0)) | ~v3_valued_0(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [cnf transformation 45] {a:0,w:20,wCS:8000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
100. v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
117. v1_funct_1(k5_relat_1(X0,X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 61] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. v3_valued_0(k5_relat_1(X0,X1)) | ~v3_valued_0(X0) | ~v1_relat_1(X0) [cnf transformation 65] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
125. 1 <=> v1_comseq_2(k5_relat_1(sK2,k2_xboole_0(sK0,sK1))) [avatar definition]
127. ~v1_comseq_2(k5_relat_1(sK2,k2_xboole_0(sK0,sK1))) <- (~1) [avatar component clause 125] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
128. ~1 [avatar split clause 79,125]
130. 2 <=> v1_comseq_2(k5_relat_1(sK2,sK1)) [avatar definition]
132. v1_comseq_2(k5_relat_1(sK2,sK1)) <- (2) [avatar component clause 130] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
133. 2 [avatar split clause 78,130]
135. 3 <=> v1_comseq_2(k5_relat_1(sK2,sK0)) [avatar definition]
137. v1_comseq_2(k5_relat_1(sK2,sK0)) <- (3) [avatar component clause 135] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
138. 3 [avatar split clause 77,135]
140. 4 <=> v3_valued_0(sK2) [avatar definition]
142. v3_valued_0(sK2) <- (4) [avatar component clause 140] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
143. 4 [avatar split clause 76,140]
145. 5 <=> v1_funct_1(sK2) [avatar definition]
147. v1_funct_1(sK2) <- (5) [avatar component clause 145] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
148. 5 [avatar split clause 75,145]
150. 6 <=> v1_relat_1(sK2) [avatar definition]
152. v1_relat_1(sK2) <- (6) [avatar component clause 150] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
153. 6 [avatar split clause 74,150]
186. v1_seq_2(k5_relat_1(sK2,sK1)) | ~v3_valued_0(k5_relat_1(sK2,sK1)) | ~v1_funct_1(k5_relat_1(sK2,sK1)) | ~v1_relat_1(k5_relat_1(sK2,sK1)) <- (2) [resolution 83,132] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
187. v1_seq_2(k5_relat_1(sK2,sK0)) | ~v3_valued_0(k5_relat_1(sK2,sK0)) | ~v1_funct_1(k5_relat_1(sK2,sK0)) | ~v1_relat_1(k5_relat_1(sK2,sK0)) <- (3) [resolution 83,137] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
189. 11 <=> v1_relat_1(k5_relat_1(sK2,sK1)) [avatar definition]
191. ~v1_relat_1(k5_relat_1(sK2,sK1)) <- (~11) [avatar component clause 189] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
193. 12 <=> v1_funct_1(k5_relat_1(sK2,sK1)) [avatar definition]
195. ~v1_funct_1(k5_relat_1(sK2,sK1)) <- (~12) [avatar component clause 193] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
197. 13 <=> v3_valued_0(k5_relat_1(sK2,sK1)) [avatar definition]
199. ~v3_valued_0(k5_relat_1(sK2,sK1)) <- (~13) [avatar component clause 197] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
201. 14 <=> v1_seq_2(k5_relat_1(sK2,sK1)) [avatar definition]
203. v1_seq_2(k5_relat_1(sK2,sK1)) <- (14) [avatar component clause 201] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
204. ~11 | ~12 | ~13 | 14 | ~2 [avatar split clause 186,130,201,197,193,189]
206. 15 <=> v1_relat_1(k5_relat_1(sK2,sK0)) [avatar definition]
208. ~v1_relat_1(k5_relat_1(sK2,sK0)) <- (~15) [avatar component clause 206] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
210. 16 <=> v1_funct_1(k5_relat_1(sK2,sK0)) [avatar definition]
212. ~v1_funct_1(k5_relat_1(sK2,sK0)) <- (~16) [avatar component clause 210] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
214. 17 <=> v3_valued_0(k5_relat_1(sK2,sK0)) [avatar definition]
216. ~v3_valued_0(k5_relat_1(sK2,sK0)) <- (~17) [avatar component clause 214] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
218. 18 <=> v1_seq_2(k5_relat_1(sK2,sK0)) [avatar definition]
220. v1_seq_2(k5_relat_1(sK2,sK0)) <- (18) [avatar component clause 218] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
221. ~15 | ~16 | ~17 | 18 | ~3 [avatar split clause 187,135,218,214,210,206]
222. v2_seq_2(k5_relat_1(sK2,sK1)) | ~v3_valued_0(k5_relat_1(sK2,sK1)) | ~v1_funct_1(k5_relat_1(sK2,sK1)) | ~v1_relat_1(k5_relat_1(sK2,sK1)) <- (2) [resolution 84,132] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
223. v2_seq_2(k5_relat_1(sK2,sK0)) | ~v3_valued_0(k5_relat_1(sK2,sK0)) | ~v1_funct_1(k5_relat_1(sK2,sK0)) | ~v1_relat_1(k5_relat_1(sK2,sK0)) <- (3) [resolution 84,137] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
225. 19 <=> v2_seq_2(k5_relat_1(sK2,sK1)) [avatar definition]
227. v2_seq_2(k5_relat_1(sK2,sK1)) <- (19) [avatar component clause 225] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
228. ~11 | ~12 | ~13 | 19 | ~2 [avatar split clause 222,130,225,197,193,189]
230. 20 <=> v2_seq_2(k5_relat_1(sK2,sK0)) [avatar definition]
232. v2_seq_2(k5_relat_1(sK2,sK0)) <- (20) [avatar component clause 230] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
233. ~15 | ~16 | ~17 | 20 | ~3 [avatar split clause 223,135,230,214,210,206]
235. ~v1_relat_1(sK2) <- (~11) [resolution 191,100] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
237. $false <- (6, ~11) [subsumption resolution 235,152] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
238. ~6 | 11 [avatar contradiction clause 237]
246. ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (~12) [resolution 195,117] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
248. ~v1_relat_1(sK2) <- (5, ~12) [subsumption resolution 246,147] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
249. $false <- (5, 6, ~12) [subsumption resolution 248,152] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
250. ~5 | ~6 | 12 [avatar contradiction clause 249]
256. ~v3_valued_0(sK2) | ~v1_relat_1(sK2) <- (~13) [resolution 199,120] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
258. ~v1_relat_1(sK2) <- (4, ~13) [subsumption resolution 256,142] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
259. $false <- (4, 6, ~13) [subsumption resolution 258,152] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
260. ~4 | ~6 | 13 [avatar contradiction clause 259]
280. ~v3_valued_0(sK2) | ~v1_relat_1(sK2) <- (~17) [resolution 216,120] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
282. ~v1_relat_1(sK2) <- (4, ~17) [subsumption resolution 280,142] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
283. $false <- (4, 6, ~17) [subsumption resolution 282,152] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
284. ~4 | ~6 | 17 [avatar contradiction clause 283]
291. ~v1_relat_1(sK2) <- (~15) [resolution 208,100] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
293. $false <- (6, ~15) [subsumption resolution 291,152] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
294. ~6 | 15 [avatar contradiction clause 293]
297. ~v1_funct_1(sK2) | ~v1_relat_1(sK2) <- (~16) [resolution 212,117] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
299. ~v1_relat_1(sK2) <- (5, ~16) [subsumption resolution 297,147] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
300. $false <- (5, 6, ~16) [subsumption resolution 299,152] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
301. ~5 | ~6 | 16 [avatar contradiction clause 300]
316. ~v2_seq_2(k5_relat_1(X0,X1)) | ~v2_seq_2(k5_relat_1(X0,X2)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | v1_comseq_2(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v1_seq_2(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v3_valued_0(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v1_funct_1(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v1_relat_1(k5_relat_1(X0,k2_xboole_0(X2,X1))) [resolution 97,88] {a:1,w:44,wCS:17600,sine:1,thAx:0,allAx:2,thDist:-2}
322. ~v2_seq_2(k5_relat_1(X0,X1)) | ~v2_seq_2(k5_relat_1(X0,X2)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | v1_comseq_2(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v1_seq_2(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v1_funct_1(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v1_relat_1(k5_relat_1(X0,k2_xboole_0(X2,X1))) [subsumption resolution 316,120] {a:1,w:38,wCS:15200,sine:1,thAx:0,allAx:2,thDist:-2}
323. ~v2_seq_2(k5_relat_1(X0,X1)) | ~v2_seq_2(k5_relat_1(X0,X2)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | v1_comseq_2(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v1_seq_2(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v1_relat_1(k5_relat_1(X0,k2_xboole_0(X2,X1))) [subsumption resolution 322,117] {a:1,w:32,wCS:12800,sine:1,thAx:0,allAx:2,thDist:-2}
324. ~v1_seq_2(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v2_seq_2(k5_relat_1(X0,X2)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | v1_comseq_2(k5_relat_1(X0,k2_xboole_0(X2,X1))) | ~v2_seq_2(k5_relat_1(X0,X1)) [subsumption resolution 323,100] {a:1,w:26,wCS:10400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
359. ~v2_seq_2(k5_relat_1(X0,X1)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | v1_comseq_2(k5_relat_1(X0,k2_xboole_0(X1,X2))) | ~v2_seq_2(k5_relat_1(X0,X2)) | ~v1_seq_2(k5_relat_1(X0,X2)) | ~v1_seq_2(k5_relat_1(X0,X1)) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [resolution 324,96] {a:2,w:34,wCS:13600,sine:1,thAx:0,allAx:3,thDist:-3}
367. v1_comseq_2(k5_relat_1(X0,k2_xboole_0(X1,X2))) | ~v3_valued_0(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v2_seq_2(k5_relat_1(X0,X1)) | ~v2_seq_2(k5_relat_1(X0,X2)) | ~v1_seq_2(k5_relat_1(X0,X2)) | ~v1_seq_2(k5_relat_1(X0,X1)) [duplicate literal removal 359] {a:2,w:28,wCS:11200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
451. ~v3_valued_0(sK2) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v2_seq_2(k5_relat_1(sK2,sK0)) | ~v2_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK0)) <- (~1) [resolution 367,127] {a:3,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
460. ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v2_seq_2(k5_relat_1(sK2,sK0)) | ~v2_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK0)) <- (~1, 4) [subsumption resolution 451,142] {a:3,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
461. ~v1_relat_1(sK2) | ~v2_seq_2(k5_relat_1(sK2,sK0)) | ~v2_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK0)) <- (~1, 4, 5) [subsumption resolution 460,147] {a:3,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
462. ~v2_seq_2(k5_relat_1(sK2,sK0)) | ~v2_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK0)) <- (~1, 4, 5, 6) [subsumption resolution 461,152] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
463. ~v2_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK0)) <- (~1, 4, 5, 6, 20) [subsumption resolution 462,232] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
464. ~v1_seq_2(k5_relat_1(sK2,sK1)) | ~v1_seq_2(k5_relat_1(sK2,sK0)) <- (~1, 4, 5, 6, 19, 20) [subsumption resolution 463,227] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
465. ~v1_seq_2(k5_relat_1(sK2,sK0)) <- (~1, 4, 5, 6, 14, 19, 20) [subsumption resolution 464,203] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
466. $false <- (~1, 4, 5, 6, 14, 18, 19, 20) [subsumption resolution 465,220] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
467. 1 | ~4 | ~5 | ~6 | ~14 | ~18 | ~19 | ~20 [avatar contradiction clause 466]
469. $false [avatar sat refutation 128,133,138,143,148,153,204,221,228,233,238,250,260,284,294,301,467] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:215,thDist:-215}
% SZS output end Proof for rfunct_1__t87_rfunct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 24
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 50
% Generated clauses: 289
% Active clauses: 68
% Passive clauses: 87
% Final active clauses: 62
% Final passive clauses: 6
% Main loop iterations started: 81
% 
% >>> Simplifying Inferences
% Duplicate literals: 56
% Fw subsumption resolutions: 49
% 
% >>> Deletion Inferences
% Simple tautologies: 27
% Forward subsumptions: 74
% 
% >>> Generating Inferences
% Binary resolution: 54
% Unit resulting resolution: 7
% Forward superposition: 61
% Backward superposition: 5
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 18
% Sat splitting refutations: 14
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 36
% SAT solver unit clauses: 8
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% unit resulting resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.003 s ( own 0.002 s ) 

