i: [1,86,0,2,1,1,0,0] t8_projred1
a: [3,150,0,2,1,86]
s: 150
i: [1,85,0,2,1,1,0,0] t8_projred1
a: [3,155,0,2,1,85]
s: 155
i: [1,89,0,4,1,1,0,0] t8_projred1
a: [3,135,0,4,1,89]
s: 135
i: [1,88,0,4,1,1,0,0] t8_projred1
a: [3,140,0,4,1,88]
s: 140
i: [1,87,0,4,1,1,0,0] t8_projred1
a: [3,145,0,4,1,87]
s: 145
i: [1,122,0,2,1,0,0,1] existence_l1_incsp_1
a: [3,160,0,2,1,122]
s: 160
i: [1,124,0,2,1,0,0,255] rc7_funct_1
a: [3,165,0,2,1,124]
s: 165
i: [1,123,0,2,1,0,0,255] rc7_funct_1
a: [3,170,0,2,1,123]
s: 170
i: [1,126,0,2,1,0,0,255] rc1_funct_1
a: [3,175,0,2,1,126]
s: 175
i: [1,125,0,2,1,0,0,255] rc1_funct_1
a: [3,180,0,2,1,125]
s: 180
i: [1,129,0,2,1,0,0,255] rc5_funct_1
a: [3,185,0,2,1,129]
s: 185
i: [1,128,0,2,1,0,0,255] rc5_funct_1
a: [3,190,0,2,1,128]
s: 190
i: [1,127,0,2,1,0,0,255] rc5_funct_1
a: [3,195,0,2,1,127]
s: 195
i: [1,131,0,2,1,0,0,255] rc2_funct_1
a: [3,200,0,2,1,131]
s: 200
i: [1,130,0,2,1,0,0,255] rc2_funct_1
a: [3,205,0,2,1,130]
s: 205
i: [1,90,0,14,4,1,0,0] t8_projred1
s: 90
i: [1,91,0,4,2,0,0,255] cc7_funct_1
s: 91
i: [1,92,0,4,2,0,0,255] cc1_funct_1
s: 92
i: [1,120,0,4,1,0,0,1] existence_m1_subset_1
s: 120
i: [1,93,0,5,2,0,0,255] t6_boole
s: 93
i: [1,96,0,5,2,0,0,1] dt_u2_incsp_1
s: 96
i: [1,97,0,5,2,0,0,1] dt_u1_incsp_1
s: 97
i: [1,119,0,5,2,0,0,255] t7_boole
s: 119
i: [1,116,0,6,2,0,0,255] t1_subset
s: 116
i: [1,121,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 121
i: [1,94,0,7,3,0,0,255] cc8_funct_1
s: 94
d: [2,208,1,5,2,1,40,94,145]
a: [3,215,1,3,1,208]
s: 215
i: [1,95,0,7,3,0,0,255] cc8_funct_1
s: 95
d: [2,210,1,5,2,1,40,94,135]
a: [3,229,1,3,1,210]
s: 229
i: [1,118,0,7,3,0,0,1] t8_boole
s: 118
i: [1,113,0,8,4,0,0,255] cc4_funct_1
s: 113
d: [2,211,1,5,2,0,40,94,120]
s: 211
d: [2,244,0,6,3,0,21,113,92]
s: 244
d: [2,246,1,4,2,1,40,244,180]
a: [3,252,1,2,1,246]
s: 252
i: [1,117,0,8,3,0,0,255] t2_subset
s: 117
i: [1,107,0,9,3,0,0,1] d7_incproj
s: 107
d: [2,273,1,7,2,1,40,117,145]
a: [3,284,1,3,1,273]
s: 284
d: [2,304,2,0,0,2,52,150,284,97]
e: 304
d: [2,305,2,0,0,2,52,215,284,91]
e: 305
d: [2,306,2,2,1,1,40,284,97]
d: [2,311,2,0,0,2,21,306,150]
e: 311
d: [2,308,2,3,1,1,40,284,91]
d: [2,313,2,0,0,2,21,308,215]
e: 313
r: 284
i: [1,98,0,14,4,0,0,1] d7_incproj
s: 98
a: [3,283,1,3,1,273]
s: 283
i: [1,99,0,14,4,0,0,1] d7_incproj
s: 99
i: [1,100,0,14,4,0,0,1] d7_incproj
s: 100
a: [3,280,1,4,1,273]
s: 280
i: [1,104,0,14,4,0,0,1] d7_incproj
s: 104
d: [2,274,1,7,2,1,40,117,140]
a: [3,289,1,4,1,274]
s: 289
i: [1,105,0,14,4,0,0,1] d7_incproj
s: 105
i: [1,106,0,14,4,0,0,1] d7_incproj
s: 106
d: [2,275,1,7,2,1,40,117,135]
a: [3,294,1,4,1,275]
s: 294
i: [1,101,0,15,4,0,0,1] d7_incproj
s: 101
a: [3,297,1,3,1,275]
s: 297
i: [1,102,0,15,4,0,0,1] d7_incproj
s: 102
i: [1,103,0,15,4,0,0,1] d7_incproj
s: 103
d: [2,327,2,4,1,1,40,280,121]
a: [3,333,2,4,1,327]
s: 333
i: [1,108,0,40,11,0,0,1] d7_incproj
s: 108
d: [2,352,2,4,1,1,40,289,121]
a: [3,358,2,4,1,352]
s: 358
d: [2,335,1,24,6,0,40,104,90]
d: [2,336,1,20,5,1,21,335,135]
d: [2,337,1,18,4,2,21,336,155]
d: [2,338,1,16,3,3,21,337,150]
a: [3,342,1,6,1,338]
s: 342
d: [2,409,2,0,0,4,52,150,155,135,342,98]
e: 409
d: [2,410,2,8,3,1,40,342,98]
d: [2,412,2,4,2,2,21,410,135]
d: [2,413,2,2,1,3,21,412,155]
d: [2,414,2,0,0,4,21,413,150]
e: 414
r: 342
a: [3,346,1,5,1,338]
s: 346
r: 346
d: [2,394,2,4,1,1,40,294,121]
a: [3,400,2,4,1,394]
s: 400
a: [3,345,1,5,1,338]
s: 345
a: [3,350,1,5,1,338]
s: 350
r: 345
r: 350
s: 346
a: [3,349,1,5,1,338]
s: 349
d: [2,426,1,3,1,2,23,345,350]
a: [3,432,1,3,1,426]
s: 432
a: [3,341,1,6,1,338]
s: 341
d: [2,416,2,12,4,1,44,104,346]
d: [2,418,2,8,3,2,21,416,135]
d: [2,419,2,6,2,3,21,418,155]
d: [2,420,2,4,1,4,21,419,150]
a: [3,424,2,4,1,420]
s: 424
d: [2,360,1,24,6,0,40,105,90]
d: [2,361,1,20,5,1,21,360,135]
d: [2,362,1,18,4,2,21,361,155]
d: [2,363,1,16,3,3,21,362,150]
a: [3,367,1,6,1,363]
s: 367
d: [2,452,2,0,0,4,52,150,155,135,367,99]
e: 452
d: [2,453,2,8,3,1,40,367,99]
d: [2,455,2,4,2,2,21,453,135]
d: [2,456,2,2,1,3,21,455,155]
d: [2,457,2,0,0,4,21,456,150]
e: 457
r: 367
a: [3,371,1,5,1,363]
s: 371
d: [2,459,2,13,4,1,44,101,371]
d: [2,462,2,8,3,2,21,459,346]
d: [2,463,2,4,2,3,21,462,135]
d: [2,464,2,2,1,4,21,463,155]
d: [2,465,2,0,0,5,21,464,150]
e: 465
r: 371
a: [3,375,1,5,1,363]
s: 375
r: 346
r: 349
r: 375
s: 350
d: [2,427,2,12,4,1,44,104,350]
d: [2,434,2,8,3,2,21,427,135]
d: [2,435,2,6,2,3,21,434,155]
d: [2,436,2,4,1,4,21,435,150]
a: [3,440,2,4,1,436]
s: 440
s: 371
s: 341
a: [3,366,1,6,1,363]
s: 366
d: [2,377,1,24,6,0,40,106,90]
d: [2,378,1,20,5,1,21,377,135]
d: [2,379,1,18,4,2,21,378,155]
d: [2,380,1,16,3,3,21,379,150]
a: [3,384,1,6,1,380]
s: 384
d: [2,490,2,0,0,4,52,150,155,135,384,100]
e: 490
d: [2,491,2,8,3,1,40,384,100]
d: [2,493,2,4,2,2,21,491,135]
d: [2,494,2,2,1,3,21,493,155]
d: [2,495,2,0,0,4,21,494,150]
e: 495
r: 384
a: [3,388,1,5,1,380]
s: 388
d: [2,498,2,13,4,1,44,102,388]
d: [2,505,2,8,3,2,21,498,371]
d: [2,506,2,4,2,3,21,505,135]
d: [2,507,2,2,1,4,21,506,155]
d: [2,508,2,0,0,5,21,507,150]
e: 508
r: 388
a: [3,392,1,5,1,380]
s: 392
d: [2,510,2,13,4,1,44,103,392]
d: [2,514,2,8,3,2,21,510,350]
d: [2,515,2,4,2,3,21,514,135]
d: [2,516,2,2,1,4,21,515,155]
d: [2,517,2,0,0,5,21,516,150]
e: 517
r: 371
r: 392
a: [3,370,1,5,1,363]
s: 370
s: 375
r: 370
d: [2,523,2,13,4,1,44,101,375]
d: [2,526,2,8,3,2,21,523,350]
d: [2,527,2,4,2,3,21,526,135]
d: [2,528,2,2,1,4,21,527,155]
d: [2,529,2,0,0,5,21,528,150]
e: 529
r: 350
s: 346
s: 349
s: 388
d: [2,538,2,13,4,1,44,103,388]
d: [2,542,2,8,3,2,21,538,346]
d: [2,543,2,4,2,3,21,542,135]
d: [2,544,2,2,1,4,21,543,155]
d: [2,545,2,0,0,5,21,544,150]
e: 545
r: 388
a: [3,387,1,5,1,380]
s: 387
s: 392
r: 387
d: [2,553,2,13,4,1,44,102,392]
d: [2,556,2,8,3,2,21,553,375]
d: [2,557,2,4,2,3,21,556,135]
d: [2,558,2,2,1,4,21,557,155]
d: [2,559,2,0,0,5,21,558,150]
e: 559
d: [2,531,2,9,3,2,21,523,135]
d: [2,532,2,7,2,3,21,531,155]
d: [2,533,2,5,1,4,21,532,150]
f: 559,545,533,517,508,495,465,457,414,380,363,338,85,86,89
% Refutation found. Thanks to Tanya!
% SZS status Theorem for projred1__t8_projred1
% SZS output start Proof for projred1__t8_projred1
1. ! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => ? [X4] : (X2 != X4 & X1 != X4 & r1_incsp_1(X0,X4,X3) & m1_subset_1(X4,u1_incsp_1(X0))))))) [input t8_projred1]
2. ~! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => ? [X4] : (X2 != X4 & X1 != X4 & r1_incsp_1(X0,X4,X3) & m1_subset_1(X4,u1_incsp_1(X0))))))) [negated conjecture 1]
11. ! [X0] : (l1_incsp_1(X0) => (v3_incproj(X0) <=> ! [X1] : (m1_subset_1(X1,u2_incsp_1(X0)) => ? [X2] : (? [X3] : (? [X4] : (r1_incsp_1(X0,X4,X1) & r1_incsp_1(X0,X3,X1) & r1_incsp_1(X0,X2,X1) & X2 != X4 & X3 != X4 & X2 != X3 & m1_subset_1(X4,u1_incsp_1(X0))) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0)))))) [input d7_incproj]
36. ~! [X0] : ((l1_incsp_1(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => ? [X4] : (X2 != X4 & X1 != X4 & r1_incsp_1(X0,X4,X3) & m1_subset_1(X4,u1_incsp_1(X0))))))) [pure predicate removal 2]
37. ~! [X0] : ((l1_incsp_1(X0) & v3_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => ? [X4] : (X2 != X4 & X1 != X4 & r1_incsp_1(X0,X4,X3) & m1_subset_1(X4,u1_incsp_1(X0))))))) [pure predicate removal 36]
38. ~! [X0] : ((l1_incsp_1(X0) & v3_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => ? [X4] : (X2 != X4 & X1 != X4 & r1_incsp_1(X0,X4,X3) & m1_subset_1(X4,u1_incsp_1(X0))))))) [pure predicate removal 37]
39. ~! [X0] : ((l1_incsp_1(X0) & v3_incproj(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => ? [X4] : (X2 != X4 & X1 != X4 & r1_incsp_1(X0,X4,X3) & m1_subset_1(X4,u1_incsp_1(X0))))))) [pure predicate removal 38]
40. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (! [X4] : (X2 = X4 | X1 = X4 | ~r1_incsp_1(X0,X4,X3) | ~m1_subset_1(X4,u1_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & (l1_incsp_1(X0) & v3_incproj(X0))) [ennf transformation 39]
41. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (! [X4] : (X2 = X4 | X1 = X4 | ~r1_incsp_1(X0,X4,X3) | ~m1_subset_1(X4,u1_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v3_incproj(X0)) [flattening 40]
48. ! [X0] : ((v3_incproj(X0) <=> ! [X1] : (? [X2] : (? [X3] : (? [X4] : (r1_incsp_1(X0,X4,X1) & r1_incsp_1(X0,X3,X1) & r1_incsp_1(X0,X2,X1) & X2 != X4 & X3 != X4 & X2 != X3 & m1_subset_1(X4,u1_incsp_1(X0))) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u2_incsp_1(X0)))) | ~l1_incsp_1(X0)) [ennf transformation 11]
61. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (! [X4] : (X2 = X4 | X1 = X4 | ~r1_incsp_1(X0,X4,X3) | ~m1_subset_1(X4,u1_incsp_1(X0))) & m1_subset_1(X3,u2_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v3_incproj(X0)) => (? [X1] : (? [X2] : (? [X3] : (! [X4] : (X2 = X4 | X1 = X4 | ~r1_incsp_1(sK0,X4,X3) | ~m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v3_incproj(sK0)) [choice axiom]
62. ? [X1] : (? [X2] : (? [X3] : (! [X4] : (X2 = X4 | X1 = X4 | ~r1_incsp_1(sK0,X4,X3) | ~m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) => (? [X2] : (? [X3] : (! [X4] : (X2 = X4 | sK1 = X4 | ~r1_incsp_1(sK0,X4,X3) | ~m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) [choice axiom]
63. ? [X2] : (? [X3] : (! [X4] : (X2 = X4 | sK1 = X4 | ~r1_incsp_1(sK0,X4,X3) | ~m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) => (? [X3] : (! [X4] : (sK2 = X4 | sK1 = X4 | ~r1_incsp_1(sK0,X4,X3) | ~m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) [choice axiom]
64. ? [X3] : (! [X4] : (sK2 = X4 | sK1 = X4 | ~r1_incsp_1(sK0,X4,X3) | ~m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(X3,u2_incsp_1(sK0))) => (! [X4] : (sK2 = X4 | sK1 = X4 | ~r1_incsp_1(sK0,X4,sK3) | ~m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(sK3,u2_incsp_1(sK0))) [choice axiom]
65. (((! [X4] : (sK2 = X4 | sK1 = X4 | ~r1_incsp_1(sK0,X4,sK3) | ~m1_subset_1(X4,u1_incsp_1(sK0))) & m1_subset_1(sK3,u2_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v3_incproj(sK0) [skolemisation 41,64,63,62,61]
66. ! [X0] : (((v3_incproj(X0) | ? [X1] : (! [X2] : (! [X3] : (! [X4] : (~r1_incsp_1(X0,X4,X1) | ~r1_incsp_1(X0,X3,X1) | ~r1_incsp_1(X0,X2,X1) | X2 = X4 | X3 = X4 | X2 = X3 | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u2_incsp_1(X0)))) & (! [X1] : (? [X2] : (? [X3] : (? [X4] : (r1_incsp_1(X0,X4,X1) & r1_incsp_1(X0,X3,X1) & r1_incsp_1(X0,X2,X1) & X2 != X4 & X3 != X4 & X2 != X3 & m1_subset_1(X4,u1_incsp_1(X0))) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u2_incsp_1(X0))) | ~v3_incproj(X0))) | ~l1_incsp_1(X0)) [nnf transformation 48]
67. ! [X0] : (((v3_incproj(X0) | ? [X1] : (! [X2] : (! [X3] : (! [X4] : (~r1_incsp_1(X0,X4,X1) | ~r1_incsp_1(X0,X3,X1) | ~r1_incsp_1(X0,X2,X1) | X2 = X4 | X3 = X4 | X2 = X3 | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u2_incsp_1(X0)))) & (! [X5] : (? [X6] : (? [X7] : (? [X8] : (r1_incsp_1(X0,X8,X5) & r1_incsp_1(X0,X7,X5) & r1_incsp_1(X0,X6,X5) & X6 != X8 & X7 != X8 & X6 != X7 & m1_subset_1(X8,u1_incsp_1(X0))) & m1_subset_1(X7,u1_incsp_1(X0))) & m1_subset_1(X6,u1_incsp_1(X0))) | ~m1_subset_1(X5,u2_incsp_1(X0))) | ~v3_incproj(X0))) | ~l1_incsp_1(X0)) [rectify 66]
68. ! [X0] : (? [X1] : (! [X2] : (! [X3] : (! [X4] : (~r1_incsp_1(X0,X4,X1) | ~r1_incsp_1(X0,X3,X1) | ~r1_incsp_1(X0,X2,X1) | X2 = X4 | X3 = X4 | X2 = X3 | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u2_incsp_1(X0))) => (! [X2] : (! [X3] : (! [X4] : (~r1_incsp_1(X0,X4,sK4(X0)) | ~r1_incsp_1(X0,X3,sK4(X0)) | ~r1_incsp_1(X0,X2,sK4(X0)) | X2 = X4 | X3 = X4 | X2 = X3 | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(sK4(X0),u2_incsp_1(X0)))) [choice axiom]
69. ! [X5,X0] : (? [X6] : (? [X7] : (? [X8] : (r1_incsp_1(X0,X8,X5) & r1_incsp_1(X0,X7,X5) & r1_incsp_1(X0,X6,X5) & X6 != X8 & X7 != X8 & X6 != X7 & m1_subset_1(X8,u1_incsp_1(X0))) & m1_subset_1(X7,u1_incsp_1(X0))) & m1_subset_1(X6,u1_incsp_1(X0))) => (? [X7] : (? [X8] : (r1_incsp_1(X0,X8,X5) & r1_incsp_1(X0,X7,X5) & r1_incsp_1(X0,sK5(X0,X5),X5) & sK5(X0,X5) != X8 & X7 != X8 & sK5(X0,X5) != X7 & m1_subset_1(X8,u1_incsp_1(X0))) & m1_subset_1(X7,u1_incsp_1(X0))) & m1_subset_1(sK5(X0,X5),u1_incsp_1(X0)))) [choice axiom]
70. ! [X5,X0] : (? [X7] : (? [X8] : (r1_incsp_1(X0,X8,X5) & r1_incsp_1(X0,X7,X5) & r1_incsp_1(X0,sK5(X0,X5),X5) & sK5(X0,X5) != X8 & X7 != X8 & sK5(X0,X5) != X7 & m1_subset_1(X8,u1_incsp_1(X0))) & m1_subset_1(X7,u1_incsp_1(X0))) => (? [X8] : (r1_incsp_1(X0,X8,X5) & r1_incsp_1(X0,sK6(X0,X5),X5) & r1_incsp_1(X0,sK5(X0,X5),X5) & sK5(X0,X5) != X8 & sK6(X0,X5) != X8 & sK5(X0,X5) != sK6(X0,X5) & m1_subset_1(X8,u1_incsp_1(X0))) & m1_subset_1(sK6(X0,X5),u1_incsp_1(X0)))) [choice axiom]
71. ! [X5,X0] : (? [X8] : (r1_incsp_1(X0,X8,X5) & r1_incsp_1(X0,sK6(X0,X5),X5) & r1_incsp_1(X0,sK5(X0,X5),X5) & sK5(X0,X5) != X8 & sK6(X0,X5) != X8 & sK5(X0,X5) != sK6(X0,X5) & m1_subset_1(X8,u1_incsp_1(X0))) => (r1_incsp_1(X0,sK7(X0,X5),X5) & r1_incsp_1(X0,sK6(X0,X5),X5) & r1_incsp_1(X0,sK5(X0,X5),X5) & sK5(X0,X5) != sK7(X0,X5) & sK6(X0,X5) != sK7(X0,X5) & sK5(X0,X5) != sK6(X0,X5) & m1_subset_1(sK7(X0,X5),u1_incsp_1(X0)))) [choice axiom]
72. ! [X0] : (((v3_incproj(X0) | (! [X2] : (! [X3] : (! [X4] : (~r1_incsp_1(X0,X4,sK4(X0)) | ~r1_incsp_1(X0,X3,sK4(X0)) | ~r1_incsp_1(X0,X2,sK4(X0)) | X2 = X4 | X3 = X4 | X2 = X3 | ~m1_subset_1(X4,u1_incsp_1(X0))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(sK4(X0),u2_incsp_1(X0)))) & (! [X5] : ((((r1_incsp_1(X0,sK7(X0,X5),X5) & r1_incsp_1(X0,sK6(X0,X5),X5) & r1_incsp_1(X0,sK5(X0,X5),X5) & sK5(X0,X5) != sK7(X0,X5) & sK6(X0,X5) != sK7(X0,X5) & sK5(X0,X5) != sK6(X0,X5) & m1_subset_1(sK7(X0,X5),u1_incsp_1(X0))) & m1_subset_1(sK6(X0,X5),u1_incsp_1(X0))) & m1_subset_1(sK5(X0,X5),u1_incsp_1(X0))) | ~m1_subset_1(X5,u2_incsp_1(X0))) | ~v3_incproj(X0))) | ~l1_incsp_1(X0)) [skolemisation 67,71,70,69,68]
85. v3_incproj(sK0) [cnf transformation 65] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
86. l1_incsp_1(sK0) [cnf transformation 65] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
89. m1_subset_1(sK3,u2_incsp_1(sK0)) [cnf transformation 65] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
90. ~r1_incsp_1(sK0,X4,sK3) | sK1 = X4 | sK2 = X4 | ~m1_subset_1(X4,u1_incsp_1(sK0)) [cnf transformation 65] {a:0,w:14,wCS:1400,nSel:1,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
98. m1_subset_1(sK5(X0,X5),u1_incsp_1(X0)) | ~m1_subset_1(X5,u2_incsp_1(X0)) | ~v3_incproj(X0) | ~l1_incsp_1(X0) [cnf transformation 72] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
99. m1_subset_1(sK6(X0,X5),u1_incsp_1(X0)) | ~m1_subset_1(X5,u2_incsp_1(X0)) | ~v3_incproj(X0) | ~l1_incsp_1(X0) [cnf transformation 72] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
100. m1_subset_1(sK7(X0,X5),u1_incsp_1(X0)) | ~m1_subset_1(X5,u2_incsp_1(X0)) | ~v3_incproj(X0) | ~l1_incsp_1(X0) [cnf transformation 72] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
101. sK5(X0,X5) != sK6(X0,X5) | ~m1_subset_1(X5,u2_incsp_1(X0)) | ~v3_incproj(X0) | ~l1_incsp_1(X0) [cnf transformation 72] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
102. sK6(X0,X5) != sK7(X0,X5) | ~m1_subset_1(X5,u2_incsp_1(X0)) | ~v3_incproj(X0) | ~l1_incsp_1(X0) [cnf transformation 72] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
103. sK5(X0,X5) != sK7(X0,X5) | ~m1_subset_1(X5,u2_incsp_1(X0)) | ~v3_incproj(X0) | ~l1_incsp_1(X0) [cnf transformation 72] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
104. r1_incsp_1(X0,sK5(X0,X5),X5) | ~m1_subset_1(X5,u2_incsp_1(X0)) | ~v3_incproj(X0) | ~l1_incsp_1(X0) [cnf transformation 72] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
105. r1_incsp_1(X0,sK6(X0,X5),X5) | ~m1_subset_1(X5,u2_incsp_1(X0)) | ~v3_incproj(X0) | ~l1_incsp_1(X0) [cnf transformation 72] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
106. r1_incsp_1(X0,sK7(X0,X5),X5) | ~m1_subset_1(X5,u2_incsp_1(X0)) | ~v3_incproj(X0) | ~l1_incsp_1(X0) [cnf transformation 72] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
133. 1 <=> m1_subset_1(sK3,u2_incsp_1(sK0)) [avatar definition]
135. m1_subset_1(sK3,u2_incsp_1(sK0)) <- (1) [avatar component clause 133] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
136. 1 [avatar split clause 89,133]
148. 4 <=> l1_incsp_1(sK0) [avatar definition]
150. l1_incsp_1(sK0) <- (4) [avatar component clause 148] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
151. 4 [avatar split clause 86,148]
153. 5 <=> v3_incproj(sK0) [avatar definition]
155. v3_incproj(sK0) <- (5) [avatar component clause 153] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
156. 5 [avatar split clause 85,153]
335. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) | sK1 = sK5(sK0,sK3) | sK2 = sK5(sK0,sK3) | ~m1_subset_1(sK5(sK0,sK3),u1_incsp_1(sK0)) [resolution 104,90] {a:1,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
336. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) | sK1 = sK5(sK0,sK3) | sK2 = sK5(sK0,sK3) | ~m1_subset_1(sK5(sK0,sK3),u1_incsp_1(sK0)) <- (1) [subsumption resolution 335,135] {a:1,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
337. ~l1_incsp_1(sK0) | sK1 = sK5(sK0,sK3) | sK2 = sK5(sK0,sK3) | ~m1_subset_1(sK5(sK0,sK3),u1_incsp_1(sK0)) <- (1, 5) [subsumption resolution 336,155] {a:1,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
338. sK1 = sK5(sK0,sK3) | sK2 = sK5(sK0,sK3) | ~m1_subset_1(sK5(sK0,sK3),u1_incsp_1(sK0)) <- (1, 4, 5) [subsumption resolution 337,150] {a:1,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
340. 33 <=> m1_subset_1(sK5(sK0,sK3),u1_incsp_1(sK0)) [avatar definition]
342. ~m1_subset_1(sK5(sK0,sK3),u1_incsp_1(sK0)) <- (~33) [avatar component clause 340] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
344. 34 <=> sK2 = sK5(sK0,sK3) [avatar definition]
346. sK2 = sK5(sK0,sK3) <- (34) [avatar component clause 344] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
348. 35 <=> sK1 = sK5(sK0,sK3) [avatar definition]
350. sK1 = sK5(sK0,sK3) <- (35) [avatar component clause 348] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
351. ~33 | 34 | 35 | ~1 | ~4 | ~5 [avatar split clause 338,153,148,133,348,344,340]
360. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) | sK1 = sK6(sK0,sK3) | sK2 = sK6(sK0,sK3) | ~m1_subset_1(sK6(sK0,sK3),u1_incsp_1(sK0)) [resolution 105,90] {a:1,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
361. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) | sK1 = sK6(sK0,sK3) | sK2 = sK6(sK0,sK3) | ~m1_subset_1(sK6(sK0,sK3),u1_incsp_1(sK0)) <- (1) [subsumption resolution 360,135] {a:1,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
362. ~l1_incsp_1(sK0) | sK1 = sK6(sK0,sK3) | sK2 = sK6(sK0,sK3) | ~m1_subset_1(sK6(sK0,sK3),u1_incsp_1(sK0)) <- (1, 5) [subsumption resolution 361,155] {a:1,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
363. sK1 = sK6(sK0,sK3) | sK2 = sK6(sK0,sK3) | ~m1_subset_1(sK6(sK0,sK3),u1_incsp_1(sK0)) <- (1, 4, 5) [subsumption resolution 362,150] {a:1,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
365. 37 <=> m1_subset_1(sK6(sK0,sK3),u1_incsp_1(sK0)) [avatar definition]
367. ~m1_subset_1(sK6(sK0,sK3),u1_incsp_1(sK0)) <- (~37) [avatar component clause 365] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
369. 38 <=> sK2 = sK6(sK0,sK3) [avatar definition]
371. sK2 = sK6(sK0,sK3) <- (38) [avatar component clause 369] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
373. 39 <=> sK1 = sK6(sK0,sK3) [avatar definition]
375. sK1 = sK6(sK0,sK3) <- (39) [avatar component clause 373] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
376. ~37 | 38 | 39 | ~1 | ~4 | ~5 [avatar split clause 363,153,148,133,373,369,365]
377. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) | sK1 = sK7(sK0,sK3) | sK2 = sK7(sK0,sK3) | ~m1_subset_1(sK7(sK0,sK3),u1_incsp_1(sK0)) [resolution 106,90] {a:1,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
378. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) | sK1 = sK7(sK0,sK3) | sK2 = sK7(sK0,sK3) | ~m1_subset_1(sK7(sK0,sK3),u1_incsp_1(sK0)) <- (1) [subsumption resolution 377,135] {a:1,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
379. ~l1_incsp_1(sK0) | sK1 = sK7(sK0,sK3) | sK2 = sK7(sK0,sK3) | ~m1_subset_1(sK7(sK0,sK3),u1_incsp_1(sK0)) <- (1, 5) [subsumption resolution 378,155] {a:1,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
380. sK1 = sK7(sK0,sK3) | sK2 = sK7(sK0,sK3) | ~m1_subset_1(sK7(sK0,sK3),u1_incsp_1(sK0)) <- (1, 4, 5) [subsumption resolution 379,150] {a:1,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
382. 40 <=> m1_subset_1(sK7(sK0,sK3),u1_incsp_1(sK0)) [avatar definition]
384. ~m1_subset_1(sK7(sK0,sK3),u1_incsp_1(sK0)) <- (~40) [avatar component clause 382] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
386. 41 <=> sK2 = sK7(sK0,sK3) [avatar definition]
388. sK2 = sK7(sK0,sK3) <- (41) [avatar component clause 386] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
390. 42 <=> sK1 = sK7(sK0,sK3) [avatar definition]
392. sK1 = sK7(sK0,sK3) <- (42) [avatar component clause 390] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
393. ~40 | 41 | 42 | ~1 | ~4 | ~5 [avatar split clause 380,153,148,133,390,386,382]
410. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (~33) [resolution 342,98] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
412. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (1, ~33) [subsumption resolution 410,135] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
413. ~l1_incsp_1(sK0) <- (1, 5, ~33) [subsumption resolution 412,155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
414. $false <- (1, 4, 5, ~33) [subsumption resolution 413,150] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
415. ~1 | ~4 | ~5 | 33 [avatar contradiction clause 414]
453. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (~37) [resolution 367,99] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
455. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (1, ~37) [subsumption resolution 453,135] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
456. ~l1_incsp_1(sK0) <- (1, 5, ~37) [subsumption resolution 455,155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
457. $false <- (1, 4, 5, ~37) [subsumption resolution 456,150] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
458. ~1 | ~4 | ~5 | 37 [avatar contradiction clause 457]
459. sK2 != sK5(sK0,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (38) [superposition 101,371] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
462. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (34, 38) [subsumption resolution 459,346] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
463. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (1, 34, 38) [subsumption resolution 462,135] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
464. ~l1_incsp_1(sK0) <- (1, 5, 34, 38) [subsumption resolution 463,155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
465. $false <- (1, 4, 5, 34, 38) [subsumption resolution 464,150] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
466. ~1 | ~4 | ~5 | ~34 | ~38 [avatar contradiction clause 465]
491. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (~40) [resolution 384,100] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
493. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (1, ~40) [subsumption resolution 491,135] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
494. ~l1_incsp_1(sK0) <- (1, 5, ~40) [subsumption resolution 493,155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
495. $false <- (1, 4, 5, ~40) [subsumption resolution 494,150] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
496. ~1 | ~4 | ~5 | 40 [avatar contradiction clause 495]
498. sK2 != sK6(sK0,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (41) [superposition 102,388] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
505. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (38, 41) [subsumption resolution 498,371] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
506. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (1, 38, 41) [subsumption resolution 505,135] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
507. ~l1_incsp_1(sK0) <- (1, 5, 38, 41) [subsumption resolution 506,155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
508. $false <- (1, 4, 5, 38, 41) [subsumption resolution 507,150] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
509. ~1 | ~4 | ~5 | ~38 | ~41 [avatar contradiction clause 508]
510. sK1 != sK5(sK0,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (42) [superposition 103,392] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
514. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (35, 42) [subsumption resolution 510,350] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
515. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (1, 35, 42) [subsumption resolution 514,135] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
516. ~l1_incsp_1(sK0) <- (1, 5, 35, 42) [subsumption resolution 515,155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
517. $false <- (1, 4, 5, 35, 42) [subsumption resolution 516,150] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
518. ~1 | ~4 | ~5 | ~35 | ~42 [avatar contradiction clause 517]
523. sK1 != sK5(sK0,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (39) [superposition 101,375] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
531. sK1 != sK5(sK0,sK3) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (1, 39) [subsumption resolution 523,135] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
532. sK1 != sK5(sK0,sK3) | ~l1_incsp_1(sK0) <- (1, 5, 39) [subsumption resolution 531,155] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
533. sK1 != sK5(sK0,sK3) <- (1, 4, 5, 39) [subsumption resolution 532,150] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
534. ~35 | ~1 | ~4 | ~5 | ~39 [avatar split clause 533,373,153,148,133,348]
538. sK2 != sK5(sK0,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (41) [superposition 103,388] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
542. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (34, 41) [subsumption resolution 538,346] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
543. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (1, 34, 41) [subsumption resolution 542,135] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
544. ~l1_incsp_1(sK0) <- (1, 5, 34, 41) [subsumption resolution 543,155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
545. $false <- (1, 4, 5, 34, 41) [subsumption resolution 544,150] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
546. ~1 | ~4 | ~5 | ~34 | ~41 [avatar contradiction clause 545]
553. sK1 != sK6(sK0,sK3) | ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (42) [superposition 102,392] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
556. ~m1_subset_1(sK3,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (39, 42) [subsumption resolution 553,375] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
557. ~v3_incproj(sK0) | ~l1_incsp_1(sK0) <- (1, 39, 42) [subsumption resolution 556,135] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
558. ~l1_incsp_1(sK0) <- (1, 5, 39, 42) [subsumption resolution 557,155] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
559. $false <- (1, 4, 5, 39, 42) [subsumption resolution 558,150] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
560. ~1 | ~4 | ~5 | ~39 | ~42 [avatar contradiction clause 559]
561. $false [avatar sat refutation 136,151,156,351,376,393,415,458,466,496,509,518,534,546,560] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:645,thDist:-645}
% SZS output end Proof for projred1__t8_projred1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% >>> Preprocessing
% Introduced skolems: 14
% Pure predicates: 6
% 
% >>> Saturation
% Initial clauses: 47
% Generated clauses: 296
% Active clauses: 78
% Passive clauses: 103
% Final active clauses: 60
% Final passive clauses: 21
% Main loop iterations started: 82
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 80
% Fw demodulations: 6
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Equational tautologies: 2
% Forward subsumptions: 39
% 
% >>> Generating Inferences
% Binary resolution: 51
% Unit resulting resolution: 5
% Forward superposition: 2
% Backward superposition: 38
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 27
% Sat splitting refutations: 16
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 58
% SAT solver unit clauses: 15
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward demodulation index maintenance: 0.001 s
% forward subsumption: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

