i: [1,236,0,2,1,1,0,0] t133_xreal_1
a: [3,383,0,2,1,236]
s: 383
i: [1,235,0,2,1,1,0,0] t133_xreal_1
a: [3,388,0,2,1,235]
s: 388
i: [1,239,0,6,2,1,0,0] t133_xreal_1
a: [3,364,0,3,1,239]
s: 364
i: [1,238,0,6,2,1,0,0] t133_xreal_1
a: [3,369,0,3,1,238]
s: 369
i: [1,237,0,5,1,1,0,0] t133_xreal_1
a: [3,378,0,5,1,237]
s: 378
i: [1,240,0,2,1,0,0,2] fc1_numbers
a: [3,393,0,2,1,240]
s: 393
i: [1,241,0,2,1,0,0,3] spc1_boole
a: [3,398,0,2,1,241]
s: 398
i: [1,242,0,2,1,0,0,2] fc1_xboole_0
a: [3,403,0,2,1,242]
s: 403
i: [1,251,0,2,1,0,0,3] fc6_ordinal1
a: [3,443,0,2,1,251]
s: 443
i: [1,250,0,2,1,0,0,3] fc6_ordinal1
a: [3,448,0,2,1,250]
s: 448
i: [1,252,0,2,1,0,0,2] spc1_numerals
a: [3,468,0,2,1,252]
s: 468
i: [1,321,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,472,0,2,1,321]
s: 472
i: [1,324,0,2,1,0,0,2] rc2_xboole_0
a: [3,480,0,2,1,324]
s: 480
i: [1,325,0,2,1,0,0,4] rc4_ordinal1
a: [3,485,0,2,1,325]
s: 485
i: [1,326,0,2,1,0,0,4] rc1_ordinal1
a: [3,490,0,2,1,326]
s: 490
i: [1,327,0,2,1,0,0,2] rc1_xcmplx_0
a: [3,495,0,2,1,327]
s: 495
i: [1,328,0,2,1,0,0,1] rc1_xreal_0
a: [3,500,0,2,1,328]
s: 500
i: [1,329,0,2,1,0,0,2] rc1_xxreal_0
a: [3,505,0,2,1,329]
s: 505
d: [2,642,0,0,0,2,21,505,472]
e: 642
r: 472
s: 505
i: [1,330,0,2,1,0,0,2] rc1_xboole_0
a: [3,510,0,2,1,330]
s: 510
i: [1,334,0,2,1,0,0,255] rc2_ordinal1
a: [3,515,0,2,1,334]
s: 515
i: [1,333,0,2,1,0,0,255] rc2_ordinal1
a: [3,520,0,2,1,333]
s: 520
i: [1,332,0,2,1,0,0,255] rc2_ordinal1
a: [3,525,0,2,1,332]
s: 525
i: [1,331,0,2,1,0,0,255] rc2_ordinal1
a: [3,530,0,2,1,331]
s: 530
i: [1,336,0,2,1,0,0,4] rc5_ordinal1
a: [3,535,0,2,1,336]
s: 535
i: [1,335,0,2,1,0,0,4] rc5_ordinal1
a: [3,540,0,2,1,335]
s: 540
i: [1,338,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,545,0,2,1,338]
s: 545
i: [1,337,0,2,1,0,0,2] rc2_xcmplx_0
a: [3,550,0,2,1,337]
s: 550
i: [1,342,0,2,1,0,0,2] rc3_xreal_0
a: [3,555,0,2,1,342]
s: 555
i: [1,341,0,2,1,0,0,2] rc3_xreal_0
a: [3,560,0,2,1,341]
s: 560
i: [1,340,0,2,1,0,0,2] rc3_xreal_0
a: [3,565,0,2,1,340]
s: 565
i: [1,339,0,2,1,0,0,2] rc3_xreal_0
a: [3,570,0,2,1,339]
s: 570
i: [1,346,0,2,1,0,0,2] rc2_xreal_0
a: [3,575,0,2,1,346]
s: 575
i: [1,345,0,2,1,0,0,2] rc2_xreal_0
a: [3,580,0,2,1,345]
s: 580
i: [1,344,0,2,1,0,0,2] rc2_xreal_0
a: [3,585,0,2,1,344]
s: 585
i: [1,343,0,2,1,0,0,2] rc2_xreal_0
a: [3,590,0,2,1,343]
s: 590
i: [1,348,0,2,1,0,0,2] rc3_xxreal_0
a: [3,595,0,2,1,348]
s: 595
i: [1,347,0,2,1,0,0,2] rc3_xxreal_0
a: [3,600,0,2,1,347]
s: 600
i: [1,350,0,2,1,0,0,2] rc2_xxreal_0
a: [3,605,0,2,1,350]
s: 605
i: [1,349,0,2,1,0,0,2] rc2_xxreal_0
a: [3,610,0,2,1,349]
s: 610
i: [1,354,0,2,1,0,0,2] rc4_xreal_0
a: [3,615,0,2,1,354]
s: 615
i: [1,353,0,2,1,0,0,2] rc4_xreal_0
a: [3,620,0,2,1,353]
s: 620
i: [1,352,0,2,1,0,0,2] rc4_xreal_0
a: [3,625,0,2,1,352]
s: 625
i: [1,351,0,2,1,0,0,2] rc4_xreal_0
a: [3,630,0,2,1,351]
s: 630
i: [1,356,0,2,1,0,0,2] rc4_xxreal_0
a: [3,635,0,2,1,356]
s: 635
i: [1,355,0,2,1,0,0,2] rc4_xxreal_0
a: [3,640,0,2,1,355]
s: 640
i: [1,307,0,3,1,0,0,255] reflexivity_r1_tarski
s: 307
i: [1,243,0,3,1,0,0,2] redefinition_k5_numbers
a: [3,408,0,3,1,243]
s: 408
i: [1,244,0,3,1,0,0,1] redefinition_k6_numbers
a: [3,413,0,3,1,244]
s: 413
r: 403
d: [2,644,0,2,1,2,23,403,413]
a: [3,648,0,2,1,644]
s: 648
i: [1,245,0,3,1,0,0,2] t1_numerals
a: [3,418,0,3,1,245]
s: 418
i: [1,246,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,423,0,3,1,246]
s: 423
i: [1,255,0,3,1,0,0,2] spc1_numerals
a: [3,453,0,3,1,255]
s: 453
i: [1,254,0,3,1,0,0,2] spc1_numerals
a: [3,458,0,3,1,254]
s: 458
d: [2,650,0,3,1,2,22,418,413]
a: [3,654,0,3,1,650]
s: 654
d: [2,656,0,3,1,2,22,458,408]
a: [3,660,0,3,1,656]
s: 660
i: [1,257,0,4,2,0,0,4] cc6_ordinal1
s: 257
i: [1,258,0,4,2,0,0,4] cc2_xcmplx_0
s: 258
d: [2,662,1,2,1,1,40,257,485]
a: [3,667,1,2,1,662]
s: 667
i: [1,259,0,4,2,0,0,4] cc2_xreal_0
s: 259
i: [1,260,0,4,2,0,0,4] cc2_xxreal_0
s: 260
d: [2,663,1,2,1,1,40,257,535]
a: [3,672,1,2,1,663]
s: 672
i: [1,261,0,4,2,0,0,255] cc1_ordinal1
s: 261
d: [2,674,1,2,1,1,40,258,485]
a: [3,679,1,2,1,674]
s: 679
i: [1,262,0,4,2,0,0,255] cc1_ordinal1
s: 262
i: [1,266,0,4,2,0,0,2] cc3_xreal_0
s: 266
d: [2,760,1,2,1,1,40,266,388]
a: [3,769,1,2,1,760]
s: 769
i: [1,267,0,4,2,0,0,2] cc4_xreal_0
s: 267
d: [2,761,1,2,1,1,40,266,383]
a: [3,774,1,2,1,761]
s: 774
i: [1,276,0,4,2,0,0,255] cc9_ordinal1
s: 276
i: [1,277,0,4,2,0,0,4] cc7_ordinal1
s: 277
d: [2,781,1,2,1,1,40,267,388]
a: [3,790,1,2,1,781]
s: 790
i: [1,278,0,4,2,0,0,4] cc3_ordinal1
s: 278
d: [2,782,1,2,1,1,40,267,383]
a: [3,795,1,2,1,782]
s: 795
i: [1,306,0,4,1,0,0,3] existence_m1_subset_1
s: 306
i: [1,247,0,4,1,0,0,2] dt_k5_numbers
a: [3,428,0,4,1,247]
s: 428
d: [2,675,1,2,1,1,40,258,535]
a: [3,684,1,2,1,675]
s: 684
i: [1,248,0,4,1,0,0,1] dt_k6_numbers
a: [3,433,0,4,1,248]
s: 433
d: [2,686,1,2,1,1,40,259,485]
a: [3,691,1,2,1,686]
s: 691
i: [1,253,0,4,1,0,0,2] spc1_numerals
a: [3,463,0,4,1,253]
s: 463
d: [2,874,0,4,1,2,22,428,408]
a: [3,878,0,4,1,874]
s: 878
d: [2,687,1,2,1,1,40,259,535]
a: [3,696,1,2,1,687]
s: 696
d: [2,880,0,4,1,2,22,433,408]
a: [3,884,0,4,1,880]
s: 884
d: [2,698,1,2,1,1,40,260,485]
a: [3,703,1,2,1,698]
s: 703
d: [2,889,0,4,1,2,22,463,408]
a: [3,893,0,4,1,889]
s: 893
i: [1,279,0,5,2,0,0,2] t6_boole
s: 279
d: [2,699,1,2,1,1,40,260,535]
a: [3,708,1,2,1,699]
s: 708
i: [1,280,0,5,2,0,0,3] cc8_ordinal1
s: 280
d: [2,710,1,2,1,1,40,261,443]
a: [3,718,1,2,1,710]
s: 718
i: [1,281,0,5,2,0,0,2] cc1_xcmplx_0
s: 281
i: [1,282,0,5,2,0,0,2] cc1_xreal_0
s: 282
d: [2,712,1,2,1,1,40,261,490]
a: [3,728,1,2,1,712]
s: 728
i: [1,303,0,5,2,0,0,255] t7_boole
s: 303
d: [2,735,1,2,1,1,40,262,443]
a: [3,743,1,2,1,735]
s: 743
i: [1,249,0,5,1,0,0,3] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,438,0,5,1,249]
s: 438
a: [3,475,0,5,2,321]
s: 475
d: [2,737,1,2,1,1,40,262,490]
a: [3,753,1,2,1,737]
s: 753
d: [2,898,0,5,2,1,22,279,413]
s: 898
d: [2,938,1,5,2,2,44,364,898]
s: 938
i: [1,287,0,6,3,0,0,255] cc2_ordinal1
s: 287
i: [1,290,0,6,3,0,0,2] cc5_xxreal_0
s: 290
d: [2,939,1,5,2,2,44,369,898]
s: 939
i: [1,292,0,6,3,0,0,2] cc5_xxreal_0
s: 292
d: [2,940,1,7,2,2,44,378,898]
s: 940
i: [1,293,0,6,3,0,0,2] cc3_xxreal_0
s: 293
i: [1,295,0,6,3,0,0,2] cc3_xxreal_0
s: 295
d: [2,975,2,8,2,3,40,940,475]
a: [3,985,2,4,1,975]
s: 985
i: [1,297,0,6,3,0,0,2] cc7_xxreal_0
s: 297
d: [2,762,1,2,1,1,40,266,500]
a: [3,779,1,2,1,762]
s: 779
i: [1,298,0,6,3,0,0,2] cc7_xxreal_0
s: 298
i: [1,299,0,6,2,0,0,255] t1_subset
s: 299
d: [2,783,1,2,1,1,40,267,500]
a: [3,800,1,2,1,783]
s: 800
i: [1,308,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 308
d: [2,802,1,2,1,1,40,276,648]
a: [3,809,1,2,1,802]
s: 809
i: [1,263,0,7,3,0,0,4] cc5_ordinal1
s: 263
i: [1,264,0,7,2,0,0,1] t2_arithm
s: 264
d: [2,803,1,2,1,1,40,276,510]
a: [3,814,1,2,1,803]
s: 814
i: [1,265,0,7,2,0,0,3] t3_arithm
s: 265
d: [2,804,1,2,1,1,40,276,630]
a: [3,819,1,2,1,804]
s: 819
i: [1,301,0,7,2,0,0,255] t3_subset
s: 301
i: [1,302,0,7,3,0,0,2] t8_boole
s: 302
d: [2,805,1,2,1,1,40,276,640]
a: [3,824,1,2,1,805]
s: 824
i: [1,256,0,8,3,0,0,255] cc10_ordinal1
s: 256
d: [2,826,1,2,1,1,40,277,648]
a: [3,833,1,2,1,826]
s: 833
i: [1,284,0,8,4,0,0,2] cc4_xxreal_0
s: 284
i: [1,286,0,8,4,0,0,2] cc6_xxreal_0
s: 286
d: [2,976,2,6,2,3,40,940,938]
a: [3,989,2,2,1,976]
s: 989
r: 989
i: [1,288,0,8,4,0,0,2] cc8_xxreal_0
s: 288
a: [3,990,2,2,1,976]
s: 990
i: [1,300,0,8,3,0,0,255] t2_subset
s: 300
i: [1,318,0,8,3,0,0,1] fc3_xcmplx_0
s: 318
d: [2,977,2,6,2,3,40,940,939]
a: [3,994,2,2,1,977]
s: 994
r: 994
i: [1,320,0,8,3,0,0,1] fc6_xreal_0
s: 320
a: [3,995,2,2,1,977]
s: 995
i: [1,305,0,9,3,0,0,255] t5_subset
s: 305
i: [1,304,0,10,3,0,0,255] t4_subset
s: 304
d: [2,1066,2,6,3,1,40,284,790]
a: [3,1081,2,2,1,1066]
s: 1081
i: [1,322,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 322
d: [2,1067,2,6,3,1,40,284,795]
a: [3,1090,2,2,1,1067]
s: 1090
i: [1,270,0,11,5,0,0,1] t2_real
s: 270
i: [1,271,0,11,5,0,0,1] t1_real
s: 271
r: 1081
r: 1090
a: [3,1080,2,2,1,1066]
s: 1080
i: [1,272,0,11,5,0,0,1] t8_real
s: 272
a: [3,1085,2,2,1,1066]
s: 1085
i: [1,273,0,11,5,0,0,1] t7_real
s: 273
i: [1,274,0,11,5,0,0,1] t3_real
s: 274
a: [3,1089,2,2,1,1067]
s: 1089
i: [1,275,0,11,5,0,0,1] t4_real
s: 275
a: [3,1094,2,2,1,1067]
s: 1094
i: [1,319,0,11,3,0,0,2] commutativity_k3_xcmplx_0
s: 319
i: [1,309,0,12,5,0,0,1] fc26_xreal_0
s: 309
d: [2,1309,1,12,4,1,40,272,378]
a: [3,1318,1,4,1,1309]
s: 1318
d: [2,1371,2,0,0,3,52,388,383,1318,320]
e: 1371
d: [2,1372,2,4,2,1,40,1318,320]
d: [2,1374,2,2,1,2,21,1372,383]
d: [2,1375,2,0,0,3,21,1374,388]
e: 1375
r: 1318
i: [1,310,0,12,5,0,0,1] fc23_xreal_0
s: 310
d: [2,1251,1,8,4,1,40,270,364]
d: [2,1265,1,6,3,2,21,1251,383]
a: [3,1268,1,2,1,1265]
s: 1268
r: 1268
i: [1,311,0,12,5,0,0,1] fc24_xreal_0
s: 311
i: [1,312,0,12,5,0,0,1] fc25_xreal_0
s: 312
a: [3,1269,1,2,1,1265]
s: 1269
i: [1,317,0,12,5,0,0,1] fc8_xcmplx_0
s: 317
d: [2,1277,1,8,4,1,40,271,364]
d: [2,1288,1,6,3,2,21,1277,1090]
d: [2,1289,1,4,2,3,21,1288,383]
a: [3,1292,1,2,1,1289]
s: 1292
i: [1,268,0,13,6,0,0,1] t5_real
s: 268
i: [1,269,0,13,6,0,0,1] t6_real
s: 269
a: [3,1317,1,4,1,1309]
s: 1317
i: [1,316,0,14,4,0,0,2] existence_m2_subset_1
s: 316
a: [3,1322,1,4,1,1309]
s: 1322
d: [2,1462,3,0,0,5,52,1080,388,1089,383,1322,312]
e: 1462
d: [2,1463,2,8,4,1,40,1322,312]
d: [2,1468,2,6,3,2,21,1463,383]
d: [2,1469,2,4,2,3,21,1468,1089]
d: [2,1470,2,2,1,4,21,1469,388]
d: [2,1471,2,0,0,5,21,1470,1080]
e: 1471
r: 1292
r: 1322
i: [1,313,0,15,5,0,0,2] dt_m2_subset_1
s: 313
i: [1,314,0,15,5,0,0,2] redefinition_m2_subset_1
s: 314
a: [3,1293,1,2,1,1289]
s: 1293
r: 939
r: 369
r: 1080
r: 1293
r: 1085
a: [3,368,0,3,1,238]
s: 368
s: 1292
a: [3,360,0,3,1,239]
s: 360
a: [3,373,0,3,1,238]
s: 373
r: 990
r: 1094
s: 989
i: [1,315,0,15,5,0,0,2] redefinition_m2_subset_1
s: 315
s: 1081
i: [1,323,0,17,4,0,0,2] spc7_arithm
s: 323
s: 1322
r: 938
r: 364
r: 368
r: 373
r: 995
r: 1081
r: 1089
s: 994
a: [3,363,0,3,1,239]
s: 363
d: [2,899,1,2,1,1,40,280,654]
d: [2,1045,2,2,1,1,40,833,260]
d: [2,1046,2,2,1,1,40,833,259]
d: [2,1278,1,8,4,1,40,271,369]
d: [2,1298,1,6,3,2,21,1278,388]
d: [2,1305,1,6,3,2,21,1277,383]
d: [2,1393,2,4,2,1,40,1268,290]
d: [2,1396,2,2,1,2,21,1393,648]
d: [2,1489,2,4,2,1,40,1293,293]
d: [2,1491,2,2,1,2,21,1489,648]
d: [2,1325,3,4,2,1,40,1085,292]
d: [2,1493,3,2,1,2,21,1325,790]
d: [2,1507,1,8,4,1,40,360,270]
d: [2,1513,1,6,3,2,21,1507,1269]
d: [2,1514,1,4,2,3,21,1513,388]
d: [2,1519,1,8,4,1,40,373,270]
d: [2,1533,1,6,3,2,21,1519,1269]
d: [2,1534,1,4,2,3,21,1533,383]
d: [2,1616,2,8,4,1,40,1322,309]
d: [2,1619,2,6,3,2,21,1616,383]
d: [2,1620,2,4,2,3,21,1619,388]
d: [2,1615,2,8,4,1,40,1322,312]
d: [2,1629,2,6,3,2,21,1615,383]
d: [2,1630,2,4,2,3,21,1629,388]
d: [2,1639,1,7,3,1,40,363,322]
d: [2,1644,1,5,2,2,21,1639,795]
f: 1644,1630,1620,1534,1514,1493,1491,1396,1375,1309,1305,1298,1046,1045,899,782,781,650,644,245,244,242,235,236,237,238,239
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t133_xreal_1
% SZS output start Proof for xreal_1__t133_xreal_1
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~(~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(k6_numbers,X0)) & ~(~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(X0,k6_numbers)) & ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X1))))) [input t133_xreal_1]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~(~r1_xxreal_0(X1,k6_numbers) & ~r1_xxreal_0(k6_numbers,X0)) & ~(~r1_xxreal_0(k6_numbers,X1) & ~r1_xxreal_0(X0,k6_numbers)) & ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X1))))) [negated conjecture 1]
10. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
11. ! [X0] : (v7_ordinal1(X0) => v1_xxreal_0(X0)) [input cc2_xxreal_0]
14. ! [X0] : ((v2_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc3_xxreal_0]
16. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
19. ! [X0] : ((v3_xxreal_0(X0) & v1_xxreal_0(X0)) => (~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0))) [input cc5_xxreal_0]
24. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
28. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
41. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
44. ! [X0,X1] : ((v1_xreal_0(X1) & ~v2_xxreal_0(X1) & v1_xreal_0(X0) & ~v2_xxreal_0(X0)) => ~v3_xxreal_0(k3_xcmplx_0(X0,X1))) [input fc25_xreal_0]
45. ! [X0,X1] : ((v1_xreal_0(X1) & ~v3_xxreal_0(X1) & v1_xreal_0(X0) & ~v3_xxreal_0(X0)) => ~v3_xxreal_0(k3_xcmplx_0(X0,X1))) [input fc26_xreal_0]
48. ! [X0,X1] : ((v1_xreal_0(X1) & v1_xreal_0(X0)) => v1_xreal_0(k3_xcmplx_0(X0,X1))) [input fc6_xreal_0]
68. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
77. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
78. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v2_xxreal_0(X0) & r1_xxreal_0(X0,X1)) => v2_xxreal_0(X1)))) [input t1_real]
81. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((v3_xxreal_0(X1) & r1_xxreal_0(X0,X1)) => v3_xxreal_0(X0)))) [input t2_real]
95. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(~v2_xxreal_0(X0) & ~v3_xxreal_0(X1) & ~r1_xxreal_0(X0,X1)))) [input t8_real]
107. ? [X0] : (? [X1] : (((r1_xxreal_0(X1,k6_numbers) | r1_xxreal_0(k6_numbers,X0)) & (r1_xxreal_0(k6_numbers,X1) | r1_xxreal_0(X0,k6_numbers)) & ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X1))) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
108. ? [X0] : (? [X1] : ((r1_xxreal_0(X1,k6_numbers) | r1_xxreal_0(k6_numbers,X0)) & (r1_xxreal_0(k6_numbers,X1) | r1_xxreal_0(X0,k6_numbers)) & ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X1)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 107]
112. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 10]
113. ! [X0] : (v1_xxreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 11]
119. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 16]
124. ! [X0] : (! [X1] : ((v3_xxreal_0(X0) | (~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 81]
125. ! [X0] : (! [X1] : (v3_xxreal_0(X0) | ~v3_xxreal_0(X1) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 124]
126. ! [X0] : (! [X1] : ((v2_xxreal_0(X1) | (~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 78]
127. ! [X0] : (! [X1] : (v2_xxreal_0(X1) | ~v2_xxreal_0(X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 126]
128. ! [X0] : (! [X1] : ((v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 95]
129. ! [X0] : (! [X1] : (v2_xxreal_0(X0) | v3_xxreal_0(X1) | r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 128]
140. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 24]
151. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v3_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 19]
152. ! [X0] : ((~v2_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 151]
153. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | (~v2_xxreal_0(X0) | ~v1_xxreal_0(X0))) [ennf transformation 14]
154. ! [X0] : ((~v3_xxreal_0(X0) & v1_xxreal_0(X0) & ~v1_xboole_0(X0)) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0)) [flattening 153]
167. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0))) [ennf transformation 45]
168. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0)) [flattening 167]
173. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0))) [ennf transformation 44]
174. ! [X0,X1] : (~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0)) [flattening 173]
187. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | (~v1_xreal_0(X1) | ~v1_xreal_0(X0))) [ennf transformation 48]
188. ! [X0,X1] : (v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0)) [flattening 187]
191. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 28]
192. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 191]
195. ? [X0] : (? [X1] : ((r1_xxreal_0(X1,k6_numbers) | r1_xxreal_0(k6_numbers,X0)) & (r1_xxreal_0(k6_numbers,X1) | r1_xxreal_0(X0,k6_numbers)) & ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(X0,X1)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : ((r1_xxreal_0(X1,k6_numbers) | r1_xxreal_0(k6_numbers,sK0)) & (r1_xxreal_0(k6_numbers,X1) | r1_xxreal_0(sK0,k6_numbers)) & ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,X1)) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
196. ? [X1] : ((r1_xxreal_0(X1,k6_numbers) | r1_xxreal_0(k6_numbers,sK0)) & (r1_xxreal_0(k6_numbers,X1) | r1_xxreal_0(sK0,k6_numbers)) & ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,X1)) & v1_xreal_0(X1)) => ((r1_xxreal_0(sK1,k6_numbers) | r1_xxreal_0(k6_numbers,sK0)) & (r1_xxreal_0(k6_numbers,sK1) | r1_xxreal_0(sK0,k6_numbers)) & ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK1)) & v1_xreal_0(sK1)) [choice axiom]
197. ((r1_xxreal_0(sK1,k6_numbers) | r1_xxreal_0(k6_numbers,sK0)) & (r1_xxreal_0(k6_numbers,sK1) | r1_xxreal_0(sK0,k6_numbers)) & ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK1)) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 108,196,195]
235. v1_xreal_0(sK0) [cnf transformation 197] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
236. v1_xreal_0(sK1) [cnf transformation 197] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
237. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK1)) [cnf transformation 197] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
238. r1_xxreal_0(k6_numbers,sK1) | r1_xxreal_0(sK0,k6_numbers) [cnf transformation 197] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
239. r1_xxreal_0(sK1,k6_numbers) | r1_xxreal_0(k6_numbers,sK0) [cnf transformation 197] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
242. v1_xboole_0(k1_xboole_0) [cnf transformation 41] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
244. k6_numbers = k1_xboole_0 [cnf transformation 68] {a:0,w:3,wCS:1200,sine:1,thAx:0,allAx:1,thDist:-1}
245. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 77] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
259. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 112] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
260. ~v7_ordinal1(X0) | v1_xxreal_0(X0) [cnf transformation 113] {a:0,w:4,wCS:1600,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
267. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 119] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
270. ~r1_xxreal_0(X0,X1) | ~v3_xxreal_0(X1) | v3_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 125] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
271. ~r1_xxreal_0(X0,X1) | ~v2_xxreal_0(X0) | v2_xxreal_0(X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 127] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
272. r1_xxreal_0(X0,X1) | v3_xxreal_0(X1) | v2_xxreal_0(X0) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 129] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
280. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 140] {a:0,w:5,wCS:2000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
290. ~v3_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 152] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
292. ~v3_xxreal_0(X0) | ~v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 152] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
293. ~v2_xxreal_0(X0) | ~v1_xboole_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 154] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
309. ~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v3_xxreal_0(X1) | ~v1_xreal_0(X0) | v3_xxreal_0(X0) [cnf transformation 168] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
312. ~v3_xxreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | v2_xxreal_0(X1) | ~v1_xreal_0(X0) | v2_xxreal_0(X0) [cnf transformation 174] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
320. v1_xreal_0(k3_xcmplx_0(X0,X1)) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 188] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
322. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 192] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
358. 1 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
360. r1_xxreal_0(k6_numbers,sK0) <- (1) [avatar component clause 358] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
362. 2 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
363. ~r1_xxreal_0(sK1,k6_numbers) <- (~2) [avatar component clause 362] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
364. r1_xxreal_0(sK1,k6_numbers) <- (2) [avatar component clause 362] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
365. 1 | 2 [avatar split clause 239,362,358]
367. 3 <=> r1_xxreal_0(sK0,k6_numbers) [avatar definition]
369. r1_xxreal_0(sK0,k6_numbers) <- (3) [avatar component clause 367] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
371. 4 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
373. r1_xxreal_0(k6_numbers,sK1) <- (4) [avatar component clause 371] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
374. 3 | 4 [avatar split clause 238,371,367]
376. 5 <=> r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK1)) [avatar definition]
378. ~r1_xxreal_0(k6_numbers,k3_xcmplx_0(sK0,sK1)) <- (~5) [avatar component clause 376] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
379. ~5 [avatar split clause 237,376]
381. 6 <=> v1_xreal_0(sK1) [avatar definition]
383. v1_xreal_0(sK1) <- (6) [avatar component clause 381] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
384. 6 [avatar split clause 236,381]
386. 7 <=> v1_xreal_0(sK0) [avatar definition]
388. v1_xreal_0(sK0) <- (7) [avatar component clause 386] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
389. 7 [avatar split clause 235,386]
401. 10 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
403. v1_xboole_0(k1_xboole_0) <- (10) [avatar component clause 401] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
404. 10 [avatar split clause 242,401]
411. 12 <=> k6_numbers = k1_xboole_0 [avatar definition]
413. k6_numbers = k1_xboole_0 <- (12) [avatar component clause 411] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
414. 12 [avatar split clause 244,411]
416. 13 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
418. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (13) [avatar component clause 416] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
419. 13 [avatar split clause 245,416]
644. v1_xboole_0(k6_numbers) <- (10, 12) [backward demodulation 403,413] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
646. 59 <=> v1_xboole_0(k6_numbers) [avatar definition]
648. v1_xboole_0(k6_numbers) <- (59) [avatar component clause 646] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
649. 59 | ~10 | ~12 [avatar split clause 644,411,401,646]
650. m1_subset_1(k6_numbers,k4_ordinal1) <- (12, 13) [forward demodulation 418,413] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
652. 60 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
654. m1_subset_1(k6_numbers,k4_ordinal1) <- (60) [avatar component clause 652] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
655. 60 | ~12 | ~13 [avatar split clause 650,416,411,652]
781. v1_xxreal_0(sK0) <- (7) [resolution 267,388] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
782. v1_xxreal_0(sK1) <- (6) [resolution 267,383] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
788. 81 <=> v1_xxreal_0(sK0) [avatar definition]
790. v1_xxreal_0(sK0) <- (81) [avatar component clause 788] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
791. 81 | ~7 [avatar split clause 781,386,788]
793. 82 <=> v1_xxreal_0(sK1) [avatar definition]
795. v1_xxreal_0(sK1) <- (82) [avatar component clause 793] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
796. 82 | ~6 [avatar split clause 782,381,793]
831. 88 <=> v7_ordinal1(k6_numbers) [avatar definition]
833. v7_ordinal1(k6_numbers) <- (88) [avatar component clause 831] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
899. v7_ordinal1(k6_numbers) <- (60) [resolution 280,654] {a:1,w:2,wCS:800,sine:3,thAx:0,allAx:2,thDist:-2}
902. 88 | ~60 [avatar split clause 899,652,831]
1045. v1_xxreal_0(k6_numbers) <- (88) [resolution 833,260] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1046. v1_xreal_0(k6_numbers) <- (88) [resolution 833,259] {a:2,w:2,wCS:800,sine:4,thAx:0,allAx:3,thDist:-3}
1051. 119 <=> v1_xxreal_0(k6_numbers) [avatar definition]
1054. 119 | ~88 [avatar split clause 1045,831,1051]
1056. 120 <=> v1_xreal_0(k6_numbers) [avatar definition]
1059. 120 | ~88 [avatar split clause 1046,831,1056]
1079. 122 <=> v2_xxreal_0(sK0) [avatar definition]
1083. 123 <=> v3_xxreal_0(sK0) [avatar definition]
1085. v3_xxreal_0(sK0) <- (123) [avatar component clause 1083] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1088. 124 <=> v2_xxreal_0(sK1) [avatar definition]
1092. 125 <=> v3_xxreal_0(sK1) [avatar definition]
1267. 151 <=> v3_xxreal_0(k6_numbers) [avatar definition]
1268. v3_xxreal_0(k6_numbers) <- (151) [avatar component clause 1267] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1269. ~v3_xxreal_0(k6_numbers) <- (~151) [avatar component clause 1267] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1277. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK1) <- (2) [resolution 271,364] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1278. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) | ~v1_xreal_0(sK0) <- (3) [resolution 271,369] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1291. 152 <=> v2_xxreal_0(k6_numbers) [avatar definition]
1293. v2_xxreal_0(k6_numbers) <- (152) [avatar component clause 1291] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1298. ~v2_xxreal_0(sK0) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (3, 7) [subsumption resolution 1278,388] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1299. ~120 | 152 | ~122 | ~3 | ~7 [avatar split clause 1298,386,367,1079,1291,1056]
1305. ~v2_xxreal_0(sK1) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k6_numbers) <- (2, 6) [subsumption resolution 1277,383] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1306. ~120 | 152 | ~124 | ~2 | ~6 [avatar split clause 1305,381,362,1088,1291,1056]
1309. v3_xxreal_0(k3_xcmplx_0(sK0,sK1)) | v2_xxreal_0(k6_numbers) | ~v1_xreal_0(k3_xcmplx_0(sK0,sK1)) | ~v1_xreal_0(k6_numbers) <- (~5) [resolution 272,378] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1316. 154 <=> v1_xreal_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
1318. ~v1_xreal_0(k3_xcmplx_0(sK0,sK1)) <- (~154) [avatar component clause 1316] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1320. 155 <=> v3_xxreal_0(k3_xcmplx_0(sK0,sK1)) [avatar definition]
1322. v3_xxreal_0(k3_xcmplx_0(sK0,sK1)) <- (155) [avatar component clause 1320] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1323. ~120 | ~154 | 152 | 155 | 5 [avatar split clause 1309,376,1320,1291,1316,1056]
1325. ~v2_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (123) [resolution 1085,292] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1372. ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) <- (~154) [resolution 1318,320] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1374. ~v1_xreal_0(sK0) <- (6, ~154) [subsumption resolution 1372,383] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1375. $false <- (6, 7, ~154) [subsumption resolution 1374,388] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1376. ~6 | ~7 | 154 [avatar contradiction clause 1375]
1393. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (151) [resolution 1268,290] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1396. ~v1_xxreal_0(k6_numbers) <- (59, 151) [subsumption resolution 1393,648] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1397. ~119 | ~59 | ~151 [avatar split clause 1396,1267,646,1051]
1489. ~v1_xboole_0(k6_numbers) | ~v1_xxreal_0(k6_numbers) <- (152) [resolution 1293,293] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1491. ~v1_xxreal_0(k6_numbers) <- (59, 152) [subsumption resolution 1489,648] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
1492. ~119 | ~59 | ~152 [avatar split clause 1491,1291,646,1051]
1493. ~v2_xxreal_0(sK0) <- (81, 123) [subsumption resolution 1325,790] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:6,thDist:-6}
1494. ~122 | ~81 | ~123 [avatar split clause 1493,1083,788,1079]
1507. ~v3_xxreal_0(sK0) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (1) [resolution 360,270] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1513. ~v3_xxreal_0(sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (1, ~151) [subsumption resolution 1507,1269] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1514. ~v3_xxreal_0(sK0) | ~v1_xreal_0(k6_numbers) <- (1, 7, ~151) [subsumption resolution 1513,388] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1515. ~120 | ~123 | ~1 | ~7 | 151 [avatar split clause 1514,1267,386,358,1083,1056]
1519. ~v3_xxreal_0(sK1) | v3_xxreal_0(k6_numbers) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (4) [resolution 373,270] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1533. ~v3_xxreal_0(sK1) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (4, ~151) [subsumption resolution 1519,1269] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1534. ~v3_xxreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (4, 6, ~151) [subsumption resolution 1533,383] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1535. ~120 | ~125 | ~4 | ~6 | 151 [avatar split clause 1534,1267,381,371,1092,1056]
1615. ~v1_xreal_0(sK1) | v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (155) [resolution 1322,312] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1616. ~v1_xreal_0(sK1) | v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (155) [resolution 1322,309] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1619. v3_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v3_xxreal_0(sK0) <- (6, 155) [subsumption resolution 1616,383] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1620. v3_xxreal_0(sK1) | v3_xxreal_0(sK0) <- (6, 7, 155) [subsumption resolution 1619,388] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1621. 123 | 125 | ~6 | ~7 | ~155 [avatar split clause 1620,1320,386,381,1092,1083]
1629. v2_xxreal_0(sK1) | ~v1_xreal_0(sK0) | v2_xxreal_0(sK0) <- (6, 155) [subsumption resolution 1615,383] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1630. v2_xxreal_0(sK1) | v2_xxreal_0(sK0) <- (6, 7, 155) [subsumption resolution 1629,388] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1631. 122 | 124 | ~6 | ~7 | ~155 [avatar split clause 1630,1320,386,381,1088,1079]
1639. r1_xxreal_0(k6_numbers,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k6_numbers) <- (~2) [resolution 363,322] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1644. r1_xxreal_0(k6_numbers,sK1) | ~v1_xxreal_0(k6_numbers) <- (~2, 82) [subsumption resolution 1639,795] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1645. ~119 | 4 | 2 | ~82 [avatar split clause 1644,793,362,371,1051]
1646. $false [avatar sat refutation 365,374,379,384,389,404,414,419,649,655,791,796,902,1054,1059,1299,1306,1323,1376,1397,1492,1494,1515,1535,1621,1631,1645] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1157,thDist:-1157}
% SZS output end Proof for xreal_1__t133_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 94
% >>> Preprocessing
% Introduced skolems: 20
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 122
% Generated clauses: 747
% Active clauses: 163
% Passive clauses: 296
% Final active clauses: 139
% Final passive clauses: 108
% Main loop iterations started: 178
% 
% >>> Simplifying Inferences
% Duplicate literals: 67
% Fw subsumption resolutions: 97
% Fw demodulations: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 31
% Forward subsumptions: 79
% 
% >>> Generating Inferences
% Binary resolution: 195
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 55
% Backward superposition: 24
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 61
% Split components: 142
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 204
% SAT solver unit clauses: 55
% SAT solver binary clauses: 76
% 
% Memory used [KB]: 5884
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% congruence closure: 0.001 s
% backward subsumption index maintenance: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

